---
id: 587d7fa7367417b2b2512bc7
title: Badilisha mitindo kulingana na data
challengeType: 6
forumTopicId: 301479
dashedName: change-styles-based-on-data
---

# --description--

D3 ni kuhusu uwasilishaji na uonyesho wa data. Inawezekana utataka kubadilisha mtindo wa vipengele kulingana na data.
Kwa mfano, unaweza kutaka kupaka rangi ya bluu kwa nukta ya data ikiwa ina thamani chini ya 20, na nyekundu vinginevyo. Unaweza kutumia kitendakazi cha mwitiko katika njia ya `style()` na kujumuisha mantiki ya masharti. Kitendakazi cha mwitiko hutumia kigezo cha `d` kuwakilisha nukta ya data:

```js
selection.style('color', d => {});
```

Njia ya `style()` haizuiliki kwa kuweka `color` tu - inaweza kutumika pia na mali nyingine za CSS.

# --instructions--

Ongeza njia ya `style()` kwenye msimbo ulioko kwenye mhariri ili kuweka `color` ya vipengele vya `h2` kwa masharti. Andika kitendakazi cha mwitiko ili ikiwa thamani ya data ni chini ya 20, irudishe nyekundu, vinginevyo irudishe kijani.

**Kumbuka:** Unaweza kutumia mantiki ya if-else, au opereta ya ternary.

# --hints--

`h2` ya kwanza inapaswa kuwa na `color` ya nyekundu.

```js
const headingTwoFirst = document.querySelector('h2');
assert.exists(headingTwoFirst); 
const color = window.getComputedStyle(headingTwoFirst)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

`h2` ya pili inapaswa kuwa na `color` ya kijani.

```js
const headingTwoSecond = document.querySelectorAll('h2')[1];
assert.exists(headingTwoSecond); 
const color = window.getComputedStyle(headingTwoSecond)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

`h2` ya tatu inapaswa kuwa na `color` ya kijani.

```js
const headingTwoThird = document.querySelectorAll('h2')[2];
assert.exists(headingTwoThird); 
const color = window.getComputedStyle(headingTwoThird)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

`h2` ya nne inapaswa kuwa na `color` ya nyekundu.

```js
const headingTwoFourth = document.querySelectorAll('h2')[3];
assert.exists(headingTwoFourth); 
const color = window.getComputedStyle(headingTwoFourth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

`h2` ya tano inapaswa kuwa na `color` ya kijani.

```js
const headingTwoFifth = document.querySelectorAll('h2')[4];
assert.exists(headingTwoFifth); 
const color = window.getComputedStyle(headingTwoFifth)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

`h2` ya sita inapaswa kuwa na `color` ya nyekundu.

```js
const headingTwoSixth = document.querySelectorAll('h2')[5];
assert.exists(headingTwoSixth); 
const color = window.getComputedStyle(headingTwoSixth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

`h2` ya saba inapaswa kuwa na `color` ya kijani.

```js
const headingTwoSeventh = document.querySelectorAll('h2')[6];
assert.exists(headingTwoSeventh); 
const color = window.getComputedStyle(headingTwoSeventh)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

`h2` ya nane inapaswa kuwa na `color` ya nyekundu.

```js
const headingTwoEighth = document.querySelectorAll('h2')[7];
assert.exists(headingTwoEighth); 
const color = window.getComputedStyle(headingTwoEighth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

`h2` ya tisa inapaswa kuwa na `color` ya nyekundu.

```js
const headingTwoNinth = document.querySelectorAll('h2')[8];
assert.exists(headingTwoNinth); 
const color = window.getComputedStyle(headingTwoNinth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD');
    // Add your code below this line

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD')
      .style('color', d => (d < 20 ? 'red' : 'green'));
  </script>
</body>
```
