---
id: 587d7faa367417b2b2512bd6
title: Ongeza kidokezo cha muda kwa kipengele cha D3
challengeType: 6
forumTopicId: 301470
dashedName: add-a-tooltip-to-a-d3-element
---

# --description--

Kidokezo cha muda kinaonyesha taarifa zaidi kuhusu kipengele kwenye ukurasa wakati mtumizi anapobeba kidole juu ya kipengele hicho. Kuna njia kadhaa za kuongeza kidokezo cha muda kwenye uonyesho wa picha. Changamoto hii inatumia kipengele cha SVG `title`.

`title` huambatana na njia ya `text()` kuongeza data kwa nguvu kwenye mabari.

# --instructions--

Ongeza kipengele cha `title` chini ya kila nodi ya `rect`. Kisha fanya wito wa njia ya `text()` ukiwa na kitendakazi cha mwitiko ili maandishi yaonyeshe thamani ya data.

# --hints--

Msimbo wako unapaswa kuwa na vipengele 9 vya `title`.

```js
assert.lengthOf(document.querySelectorAll('title'), 9);
```

Kipengele cha kwanza cha `title` kinapaswa kuwa na maandishi ya kidokezo cha muda ya `12`.

```js
assert.strictEqual(document.querySelectorAll('title')[0]?.textContent, '12');
```

Kipengele cha pili cha `title` kinapaswa kuwa na maandishi ya kidokezo cha muda ya `31`.

```js
assert.strictEqual(document.querySelectorAll('title')[1]?.textContent, '31');
```

Kipengele cha tatu cha `title` kinapaswa kuwa na maandishi ya kidokezo cha muda ya `22`.

```js
assert.strictEqual(document.querySelectorAll('title')[2]?.textContent, '22');
```

Kipengele cha nne cha `title` kinapaswa kuwa na maandishi ya kidokezo cha muda ya `17`.

```js
assert.strictEqual(document.querySelectorAll('title')[3]?.textContent, '17');
```

Kipengele cha tano cha `title` kinapaswa kuwa na maandishi ya kidokezo cha muda ya `25`.

```js
assert.strictEqual(document.querySelectorAll('title')[4]?.textContent, '25');
```

Kipengele cha sita cha `title` kinapaswa kuwa na maandishi ya kidokezo cha muda ya `18`.

```js
assert.strictEqual(document.querySelectorAll('title')[5]?.textContent, '18');
```

Kipengele cha saba cha `title` kinapaswa kuwa na maandishi ya kidokezo cha muda ya `29`.

```js
assert.strictEqual(document.querySelectorAll('title')[6]?.textContent, '29');
```

Kipengele cha nane cha `title` kinapaswa kuwa na maandishi ya kidokezo cha muda ya `14`.

```js
assert.strictEqual(document.querySelectorAll('title')[7]?.textContent, '14');
```

Kipengele cha tisa cha `title` kinapaswa kuwa na maandishi ya kidokezo cha muda ya `9`.

```js
assert.strictEqual(document.querySelectorAll('title')[8]?.textContent, '9');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy')
      .attr('class', 'bar');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - (d * 3 + 3));
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy')
      .attr('class', 'bar')
      .append('title')
      .text(d => d);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - (d * 3 + 3));
  </script>
</body>
```
