---
id: 587d7fac367417b2b2512bdb
title: Weka kikoa na safu ya thamani kwenye kipimo
challengeType: 6
forumTopicId: 301491
dashedName: set-a-domain-and-a-range-on-a-scale
---

# --description--

Kwa chaguo-msingi, vipimo hutumia uhusiano wa utambulisho. Hii ina maana thamani ya ingizo inaendana na thamani ya matokeo. Hata hivyo, vipimo vinaweza kuwa na ufanisi zaidi na kuvutia zaidi.

Sema seti ya data ina thamani kuanzia 50 hadi 480. Hii ni taarifa ya ingizo kwa kipimo, pia inajulikana kama <dfn>kikoa</dfn>.

Unataka kupangilia pointi hizo kando ya mhimili wa `x` kwenye SVG, kati ya vitengo 10 na vitengo 500. Hii ni taarifa ya matokeo, pia inajulikana kama <dfn>safu ya thamani</dfn>.

Njia za `domain()` na `range()` zinaweka thamani hizi kwa kipimo. Njia zote mbili zinachukua safu ya data yenye vipengele viwili angalau kama hoja. Hapa kuna mfano:

```js
scale.domain([50, 480]);
scale.range([10, 500]);
scale(50);
scale(480);
scale(325);
scale(750);
d3.scaleLinear();
```

Kwa mpangilio, thamani zifuatazo zitaonyeshwa kwenye konsoli: `10`, `500`, `323.37`, na `807.67`.

Tambua kuwa kipimo kinatumia uhusiano wa mstari kati ya thamani za kikoa na safu ya thamani ili kubaini matokeo yanayopaswa kwa nambari fulani. Thamani ya kiwango cha chini katika kikoa (50) inaendana na kiwango cha chini (10) katika safu ya thamani.

# --instructions--

Tengeneza kipimo na weka kikoa chake kuwa `[250, 500]` na safu ya thamani kuwa `[10, 150]`.

**Kumbuka:** Unaweza kuunganisha njia za `domain()` na `range()` kwenye mabadiliko ya `scale`.

# --hints--

Msimbo wako unapaswa kutumia njia ya `domain()`.

```js
assert.match(code, /\.domain/g);
```

`domain()` ya `scale` inapaswa kuwekwa kuwa `[250, 500]`.

```js
assert.deepEqual(scale.domain(), [250, 500]);
```

Msimbo wako unapaswa kutumia njia ya `range()`.

```js
assert.match(code, /\.range/g);
```

`range()` ya `scale` inapaswa kuwekwa kuwa `[10, 150]`.

```js
assert.deepEqual(scale.range(), [10, 150]);
```

Maandishi katika `h2` yanapaswa kuwa `-102`.

```js
assert.strictEqual(document.querySelector('h2')?.textContent, '-102');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    // Add your code below this line
    const scale = d3.scaleLinear();



    // Add your code above this line
    const output = scale(50);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const scale = d3.scaleLinear();
    scale.domain([250, 500]);
    scale.range([10, 150]);
    const output = scale(50);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```
