---
id: 587d7fa9367417b2b2512bcf
title: Badilisha kwa nguvu urefu wa kila mstari
challengeType: 6
forumTopicId: 301486
dashedName: dynamically-change-the-height-of-each-bar
---

# --description--

Urefu wa kila mstari unaweza kuwekwa kwa thamani ya kipengele cha data katika safu ya data, kama vile jinsi thamani ya `x` ilivyowekwa kwa nguvu.

```js
selection.attr('property', (d, i) => {});
```

Hapa `d` itakuwa thamani ya kipengele cha data, na `i` itakuwa kielezo cha kipengele cha data katika safu ya data.

# --instructions--

Badilisha kitendakazi cha mwitiko cha sifa ya `height` ili kurudisha thamani ya data ikizidishwa kwa 3.

**Kumbuka:** Kumbuka kwamba kuzidisha vipengele vyote vya data kwa nambari moja huongeza ukubwa wa data (kama vile kuongezea picha). Hii husaidia kuona tofauti kati ya thamani za mistari katika mfano huu.

# --hints--

`rect` ya kwanza inapaswa kuwa na `height` ya `36`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('height'),
  '36'
);
```

`rect` ya pili inapaswa kuwa na `height` ya `93`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('height'),
  '93'
);
```

`rect` ya tatu inapaswa kuwa na `height` ya `66`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('height'),
  '66'
);
```

`rect` ya nne inapaswa kuwa na `height` ya `51`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('height'),
  '51'
);
```

`rect` ya tano inapaswa kuwa na `height` ya `75`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('height'),
  '75'
);
```

`rect` ya sita inapaswa kuwa na `height` ya `54`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('height'),
  '54'
);
```

`rect` ya saba inapaswa kuwa na `height` ya `87`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('height'),
  '87'
);
```

`rect` ya nane inapaswa kuwa na `height` ya `42`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('height'),
  '42'
);
```

`rect` ya tisa inapaswa kuwa na `height` ya `27`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('height'),
  '27'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        // Add your code below this line


        
        
        // Add your code above this line
      });
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        return d * 3;
      });
  </script>
</body>
```
