---
id: a7f4d8f2483413a6ce226cac
title: محول الأرقام الرومانية
challengeType: 5
saveSubmissionToDB: true
forumTopicId: 16044
dashedName: roman-numeral-converter
---

# --description--

حوّل الرقم المعطى إلى رقم روماني.

| الأرقام الرومانية | الأرقام العربية |
|----------------|-----------------|
| M              | 1000            |
| CM             | 900             |
| D              | 500             |
| CD             | 400             |
| C              | 100             |
| XC             | 90              |
| L              | 50              |
| XL             | 40              |
| X              | 10              |
| IX             | 9               |
| V              | 5               |
| IV             | 4               |
| I              | 1               |

يجب تقديم جميع الأجوبة بالأرقام الرومانية بحروف كبيرة.

# --hints--

`convertToRoman(2)` يجب أن تُرجع السلسلة النصية `II`.

```js
assert.deepEqual(convertToRoman(2), 'II');
```

`convertToRoman(3)` يجب أن تُرجع السلسلة النصية `III`.

```js
assert.deepEqual(convertToRoman(3), 'III');
```

`convertToRoman(4)` يجب أن تُرجع السلسلة النصية `IV`.

```js
assert.deepEqual(convertToRoman(4), 'IV');
```

`convertToRoman(5)` يجب أن تُرجع السلسلة النصية `V`.

```js
assert.deepEqual(convertToRoman(5), 'V');
```

`convertToRoman(9)` يجب أن تُرجع السلسلة النصية `IX`.

```js
assert.deepEqual(convertToRoman(9), 'IX');
```

`convertToRoman(12)` يجب أن تُرجع السلسلة النصية `XII`.

```js
assert.deepEqual(convertToRoman(12), 'XII');
```

`convertToRoman(16)` يجب أن تُرجع السلسلة النصية `XVI`.

```js
assert.deepEqual(convertToRoman(16), 'XVI');
```

`convertToRoman(29)` يجب أن تُرجع السلسلة النصية `XXIX`.

```js
assert.deepEqual(convertToRoman(29), 'XXIX');
```

`convertToRoman(44)` يجب أن تُرجع السلسلة النصية `XLIV`.

```js
assert.deepEqual(convertToRoman(44), 'XLIV');
```

`convertToRoman(45)` يجب أن تُرجع السلسلة النصية `XLV`.

```js
assert.deepEqual(convertToRoman(45), 'XLV');
```

`convertToRoman(68)` يجب أن تُرجع السلسلة النصية `LXVIII`

```js
assert.deepEqual(convertToRoman(68), 'LXVIII');
```

`convertToRoman(83)` يجب أن تُرجع السلسلة النصية `LXXXIII`

```js
assert.deepEqual(convertToRoman(83), 'LXXXIII');
```

`convertToRoman(97)` يجب أن تُرجع السلسلة النصية `XCVII`

```js
assert.deepEqual(convertToRoman(97), 'XCVII');
```

`convertToRoman(99)` يجب أن تُرجع السلسلة النصية `XCIX`

```js
assert.deepEqual(convertToRoman(99), 'XCIX');
```

`convertToRoman(400)` يجب أن تُرجع السلسلة النصية `CD`

```js
assert.deepEqual(convertToRoman(400), 'CD');
```

`convertToRoman(500)` يجب أن تُرجع السلسلة النصية `D`

```js
assert.deepEqual(convertToRoman(500), 'D');
```

`convertToRoman(501)` يجب أن تُرجع السلسلة النصية `DI`

```js
assert.deepEqual(convertToRoman(501), 'DI');
```

`convertToRoman(649)` يجب أن تُرجع السلسلة النصية `DCXLIX`

```js
assert.deepEqual(convertToRoman(649), 'DCXLIX');
```

`convertToRoman(798)` يجب أن تُرجع السلسلة النصية `DCCXCVIII`

```js
assert.deepEqual(convertToRoman(798), 'DCCXCVIII');
```

`convertToRoman(891)` يجب أن تُرجع السلسلة النصية `DCCCXCI`

```js
assert.deepEqual(convertToRoman(891), 'DCCCXCI');
```

`convertToRoman(1000)` يجب أن تُرجع السلسلة النصية `M`

```js
assert.deepEqual(convertToRoman(1000), 'M');
```

`convertToRoman(1004)` يجب أن تُرجع السلسلة النصية `MIV`

```js
assert.deepEqual(convertToRoman(1004), 'MIV');
```

`convertToRoman(1006)` يجب أن تُرجع السلسلة النصية `MVI`

```js
assert.deepEqual(convertToRoman(1006), 'MVI');
```

`convertToRoman(1023)` يجب أن تُرجع السلسلة النصية `MXXIII`

```js
assert.deepEqual(convertToRoman(1023), 'MXXIII');
```

`convertToRoman(2014)` يجب أن تُرجع السلسلة النصية `MMXIV`

```js
assert.deepEqual(convertToRoman(2014), 'MMXIV');
```

`convertToRoman(3999)` يجب أن تُرجع السلسلة النصية `MMMCMXCIX`

```js
assert.deepEqual(convertToRoman(3999), 'MMMCMXCIX');
```

# --seed--

## --seed-contents--

```js
function convertToRoman(num) {
 return num;
}

convertToRoman(36);
```

# --solutions--

```js
function convertToRoman(num) {
  var ref = [['M', 1000], ['CM', 900], ['D', 500], ['CD', 400], ['C', 100], ['XC', 90], ['L', 50], ['XL', 40], ['X', 10], ['IX', 9], ['V', 5], ['IV', 4], ['I', 1]];
  var res = [];
  ref.forEach(function(p) {
    while (num >= p[1]) {
      res.push(p[0]);
      num -= p[1];
    }
  });
  return res.join('');
}
```
