---
id: aaa48de84e1ecc7c742e1124
title: مدقق الكلمات المتناظرة
challengeType: 5
saveSubmissionToDB: true
forumTopicId: 16004
dashedName: palindrome-checker
---

# --description--

أرجع `true` إذا كانت السلسلة النصية المعطاة كلمة متناظرة. وإلا، أرجع `false`.

<dfn>الكلمة المتناظرة</dfn> هي كلمة أو جملة تُكتب بنفس الطريقة من الأمام والخلف، متجاهلة علامات الترقيم والحالة والمسافات.

**ملاحظة:** يجب إزالة **كل الأحرف غير الأبجدية الرقمية** (علامات الترقيم، الفراغات والرموز) وتحويل كل شيء إلى نفس الحالة (حروف صغيرة أو كبيرة) للتحقق من الكلمات المتناظرة.

سوف نمرر سلاسل نصية بصيغ مختلفة، مثل `racecar`، و`RaceCar`، و`race CAR` من بين أمثلة أخرى.

سوف نمرر أيضًا سلاسل نصية تحتوي على رموز خاصة، مثل `2A3*3a2`، و`2A3 3a2`، و`2_A3*3#A2`.

# --hints--

`palindrome("eye")` يجب أن تُرجع قيمة منطقية.

```js
assert(typeof palindrome('eye') === 'boolean');
```

`palindrome("eye")` يجب أن تُرجع `true`.

```js
assert(palindrome('eye') === true);
```

`palindrome("_eye")` يجب أن تُرجع `true`.

```js
assert(palindrome('_eye') === true);
```

`palindrome("race car")` يجب أن تُرجع `true`.

```js
assert(palindrome('race car') === true);
```

`palindrome("not a palindrome")` يجب أن تُرجع `false`.

```js
assert(palindrome('not a palindrome') === false);
```

`palindrome("A man, a plan, a canal. Panama")` يجب أن تُرجع `true`.

```js
assert(palindrome('A man, a plan, a canal. Panama') === true);
```

`palindrome("never odd or even")` يجب أن تُرجع `true`.

```js
assert(palindrome('never odd or even') === true);
```

`palindrome("nope")` يجب أن تُرجع `false`.

```js
assert(palindrome('nope') === false);
```

`palindrome("almostomla")` يجب أن تُرجع `false`.

```js
assert(palindrome('almostomla') === false);
```

`palindrome("My age is 0, 0 si ega ym.")` يجب أن تُرجع `true`.

```js
assert(palindrome('My age is 0, 0 si ega ym.') === true);
```

`palindrome("1 eye for of 1 eye.")` يجب أن تُرجع `false`.

```js
assert(palindrome('1 eye for of 1 eye.') === false);
```

`palindrome("0_0 (: /-\ :) 0-0")` يجب أن تُرجع `true`.

```js
assert(palindrome('0_0 (: /- :) 0-0') === true);
```

`palindrome("five|\_/|four")` يجب أن تُرجع `false`.

```js
assert(palindrome('five|_/|four') === false);
```

# --seed--

## --seed-contents--

```js
function palindrome(str) {
  return true;
}

palindrome("eye");
```

# --solutions--

```js
function palindrome(str) {
  var string = str.toLowerCase().split(/[^A-Za-z0-9]/gi).join('');
  var aux = string.split('');
  if (aux.join('') === aux.reverse().join('')){
    return true;
  }

  return false;
}
```
