---
id: aff0395860f5d3034dc0bfc9
title: مدقق رقم الهاتف
challengeType: 5
saveSubmissionToDB: true
forumTopicId: 16090
dashedName: telephone-number-validator
---

# --description--

أرجع `true` إذا بدت السلسلة النصية الممررة كرقم هاتف أمريكي صالح.

يمكن للمستخدم ملء حقل النموذج بأي طريقة يختارها طالما أن التنسيق مطابق لرقم أمريكي صالح. فيما يلي أمثلة على تنسيقات صالحة لأرقام الولايات المتحدة (راجع الاختبارات أدناه للمتغيرات الأخرى):

<blockquote>555-555-5555<br>(555)555-5555<br>(555) 555-5555<br>555 555 5555<br>5555555555<br>1 555 555 5555</blockquote>

في هذا التحدٍّ ستُعرض عليك سلسلة نصية مثل `800-692-7753` أو `8oo-six427676;laskdjf`. مهمتك هي التحقق من صحة رقم الهاتف الأمريكي أو رفضه بناءً على أي تركيبة من التنسيقات المقدمة أعلاه. رمز المنطقة مطلوب. إذا تم توفير رمز البلد، يجب التأكد من أن رمز البلد هو `1`. أرجع `true` إذا كانت السلسلة نصية تمثل رقم هاتف أمريكي صالح؛ وإلا أرجع `false`.

# --hints--

يجب أن تُرجع `telephoneCheck("555-555-5555")` قيمة منطقية.

```js
assert(typeof telephoneCheck('555-555-5555') === 'boolean');
```

يجب أن تُرجع `telephoneCheck("1 555-555-5555")` القيمة `true`.

```js
assert(telephoneCheck('1 555-555-5555') === true);
```

يجب أن تُرجع `telephoneCheck("1 (555) 555-5555")` القيمة `true`.

```js
assert(telephoneCheck('1 (555) 555-5555') === true);
```

يجب أن تُرجع `telephoneCheck("5555555555")` القيمة `true`.

```js
assert(telephoneCheck('5555555555') === true);
```

يجب أن تُرجع `telephoneCheck("555-555-5555")` القيمة `true`.

```js
assert(telephoneCheck('555-555-5555') === true);
```

يجب أن تُرجع `telephoneCheck("(555)555-5555")` القيمة `true`.

```js
assert(telephoneCheck('(555)555-5555') === true);
```

يجب أن تُرجع `telephoneCheck("1(555)555-5555")` القيمة `true`.

```js
assert(telephoneCheck('1(555)555-5555') === true);
```

يجب أن تُرجع `telephoneCheck("555-5555")` القيمة `false`.

```js
assert(telephoneCheck('555-5555') === false);
```

يجب أن تُرجع `telephoneCheck("5555555")` القيمة `false`.

```js
assert(telephoneCheck('5555555') === false);
```

يجب أن تُرجع `telephoneCheck("1 555)555-5555")` القيمة `false`.

```js
assert(telephoneCheck('1 555)555-5555') === false);
```

يجب أن تُرجع `telephoneCheck("1 555 555 5555")` القيمة `true`.

```js
assert(telephoneCheck('1 555 555 5555') === true);
```

يجب أن تُرجع `telephoneCheck("1 456 789 4444")` القيمة `true`.

```js
assert(telephoneCheck('1 456 789 4444') === true);
```

يجب أن تُرجع `telephoneCheck("123**&!!asdf#")` القيمة `false`.

```js
assert(telephoneCheck('123**&!!asdf#') === false);
```

يجب أن تُرجع `telephoneCheck("55555555")` القيمة `false`.

```js
assert(telephoneCheck('55555555') === false);
```

يجب أن تُرجع `telephoneCheck("(6054756961)")` القيمة `false`.

```js
assert(telephoneCheck('(6054756961)') === false);
```

يجب أن تُرجع `telephoneCheck("2 (757) 622-7382")` القيمة `false`.

```js
assert(telephoneCheck('2 (757) 622-7382') === false);
```

يجب أن تُرجع `telephoneCheck("0 (757) 622-7382")` القيمة `false`.

```js
assert(telephoneCheck('0 (757) 622-7382') === false);
```

يجب أن تُرجع `telephoneCheck("-1 (757) 622-7382")` القيمة `false`.

```js
assert(telephoneCheck('-1 (757) 622-7382') === false);
```

يجب أن تُرجع `telephoneCheck("2 757 622-7382")` القيمة `false`.

```js
assert(telephoneCheck('2 757 622-7382') === false);
```

يجب أن تُرجع `telephoneCheck("10 (757) 622-7382")` القيمة `false`.

```js
assert(telephoneCheck('10 (757) 622-7382') === false);
```

يجب أن تُرجع `telephoneCheck("27576227382")` القيمة `false`.

```js
assert(telephoneCheck('27576227382') === false);
```

يجب أن تُرجع `telephoneCheck("(275)76227382")` القيمة `false`.

```js
assert(telephoneCheck('(275)76227382') === false);
```

يجب أن تُرجع `telephoneCheck("2(757)6227382")` القيمة `false`.

```js
assert(telephoneCheck('2(757)6227382') === false);
```

يجب أن تُرجع `telephoneCheck("2(757)622-7382")` القيمة `false`.

```js
assert(telephoneCheck('2(757)622-7382') === false);
```

يجب أن تُرجع `telephoneCheck("555)-555-5555")` القيمة `false`.

```js
assert(telephoneCheck('555)-555-5555') === false);
```

يجب أن تُرجع `telephoneCheck("(555-555-5555")` القيمة `false`.

```js
assert(telephoneCheck('(555-555-5555') === false);
```

يجب أن تُرجع `telephoneCheck("(555)5(55?)-5555")` القيمة `false`.

```js
assert(telephoneCheck('(555)5(55?)-5555') === false);
```

يجب أن تُرجع `telephoneCheck("55 55-55-555-5")` القيمة `false`.

```js
assert(telephoneCheck('55 55-55-555-5') === false);
```

يجب أن تُرجع `telephoneCheck("11 555-555-5555")` القيمة `false`.

```js
assert(telephoneCheck('11 555-555-5555') === false);
```

عند استدعاء `telephoneCheck()` مع أي رقم صالح، يجب أن تُرجع `true`.

```js

const validPatterns = [
  '1 XXX-XXX-XXXX',
  '1 (XXX) XXX-XXXX',
  '1(XXX)XXX-XXXX',
  '1 XXX XXX XXXX',
  'XXXXXXXXXX',
  'XXX-XXX-XXXX',
  '(XXX)XXX-XXXX',
];

validPatterns.forEach(pattern => {
  while (pattern.includes('X')) {
    pattern = pattern.replace('X',  Math.floor(Math.random() * 7) + 2); //While this may seem weird at first, it's required for the CI build to pass
    //This is apparently because the solution provided for CI purposes actually checks for valid area and exchange codes.
  }
  assert.isTrue(telephoneCheck(pattern));
});
```

عند استدعاء `telephoneCheck()` مع رقم غير صالح، يجب أن تُرجع `false`.

```js

const invalidPatterns = [
  '10 XXX-XXX-XXXX',
  '1 (XX)XXX-XXXX',
  '1!(XXX)XXX-XXXX',
  '-1 XXX XXX XXXX',
  'XXXXXXXX',
  'XXX#XXX-XXXX',
  '(XXXXXX-XXXX',
];

invalidPatterns.forEach(pattern => {
  while (pattern.includes('X')) {
    pattern = pattern.replace('X',  Math.floor(Math.random() * 10));
  }
  assert.isFalse(telephoneCheck(pattern));
});
```


# --seed--

## --seed-contents--

```js
function telephoneCheck(str) {
  return true;
}

telephoneCheck("555-555-5555");
```

# --solutions--

```js
var re = /^([+]?1[\s]?)?((?:[(](?:[2-9]1[02-9]|[2-9][02-8][0-9])[)][\s]?)|(?:(?:[2-9]1[02-9]|[2-9][02-8][0-9])[\s.-]?)){1}([2-9]1[02-9]|[2-9][02-9]1|[2-9][02-9]{2}[\s.-]?){1}([0-9]{4}){1}$/;

function telephoneCheck(str) {
  return re.test(str);
}

telephoneCheck("555-555-5555");
```
