---
id: 68da6138d70960c197ec5eee
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Тепер ви налаштуєте структуру даних для збереження своїх результатів. Створіть змінну з іменем `result` і ініціалізуйте її як порожній список. Цей список зберігатиме всі дійсні комбінації круглих дужок, які ви згенеруєте.

Оновіть оператор return, щоб повертати `result` замість порожнього списку.

# --hints--

Ви повинні оголосити змінну з іменем `result` у вашій функції `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("result")
`) })
```

Ви повинні ініціалізувати `result` як порожній список.

```js
({ test: () => runPython(`
_result = _Node(_code).find_function("gen_parentheses").find_variable("result")
assert _result.is_equivalent("result = []") or _result.is_equivalent("result = list()")
`) })
```

Ви повинні повертати `result` у кінці вашої функції.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_return("result")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    return []
--fcc-editable-region--
```
