---
id: 68da6138d70960c197ec5ef2
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Якщо поточний рядок ще не завершений, потрібно дослідити наступні можливі стани. Додайте умову `else` до вашої інструкції `if`.

Всередині блоку `else` ви будете обробляти додавання відкриваючих круглих дужок. Додайте інструкцію `if`, щоб перевірити, чи `opens_used < pairs`. Це гарантує, що ви не використаєте більше відкриваючих дужок, ніж дозволено.

Якщо ця умова істинна, додайте новий кортеж до `queue`: `(current + '(', opens_used + 1, closes_used)`. Це представляє стан після додавання відкриваючої дужки.

# --hints--

У вас має бути умова `else` для вашої інструкції `if` всередині циклу `while`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse
`) })
```

У вас має бути інструкція `if` всередині блоку `else`, яка перевіряє, чи `opens_used < pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1].is_equivalent("opens_used < pairs")
`) })
```

Ви маєте додати `(current + '(', opens_used + 1, closes_used)` до черги, коли умова істинна.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[1].has_call("queue.append((current + '(', opens_used + 1, closes_used))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        if len(current) == 2 * pairs:
            result.append(current)
        
--fcc-editable-region--
    return result

print(gen_parentheses(1))
```
