---
id: 68da6138d70960c197ec5ef3
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Тепер потрібно опрацювати додавання закриваючих круглих дужок. Головне правило — ви можете додати закриваючу дужку лише якщо це зберігає баланс, тобто має бути більше відкриваючих дужок, ніж закриваючих.

Все ще всередині блоку `else` додайте другу інструкцію `if`, щоб перевірити, чи `closes_used < opens_used`.

Якщо ця умова істинна, додайте ще один новий кортеж до `queue`: `(current + ')', opens_used, closes_used + 1)`. Це представляє стан після додавання закриваючої дужки.

# --hints--

У вас має бути друга інструкція `if` всередині блоку `else`.

```js
({ test: () => runPython(`
import ast
assert isinstance(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1], ast.If)
`) })
```

Ваша друга вкладена інструкція `if` має перевіряти, чи `closes_used < opens_used`.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1].test).is_equivalent("closes_used < opens_used")
`) })
```

Ви маєте додати `(current + ')', opens_used, closes_used + 1)` до черги, коли умова істинна.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1]).has_call("queue.append((current + ')', opens_used, closes_used + 1))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        
        if len(current) == 2 * pairs:
            result.append(current)
--fcc-editable-region--
        else:
            if opens_used < pairs:
                queue.append((current + '(', opens_used + 1, closes_used))
            
--fcc-editable-region--
    return result

print(gen_parentheses(1))
```
