---
id: 68da6138d70960c197ec5eef
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Для підходу пошуку у ширину ви будете використовувати чергу, щоб відстежувати різні стани під час побудови комбінацій круглих дужок. Кожен стан буде представлений як кортеж, що містить три елементи:

- Поточний рядок, який будується
- Кількість відкриваючих круглих дужок, використаних до цього моменту
- Кількість закриваючих круглих дужок, використаних до цього моменту

Створіть змінну з назвою `queue` і ініціалізуйте її списком, що містить один кортеж: `('', 0, 0)`. Це представляє початковий стан з порожнім рядком і нульовою кількістю використаних дужок.

# --hints--

Ви повинні оголосити змінну з назвою `queue` у вашій функції `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("queue")
`) })
```

Ви повинні ініціалізувати `queue` списком, що містить кортеж `('', 0, 0)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("queue").is_equivalent("queue = [('', 0, 0)]")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    result = []
    return result
--fcc-editable-region--
```
