---
id: 68da6138d70960c197ec5eed
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Далі потрібно перевірити, що число пар принаймні одне, адже не можна згенерувати комбінації круглих дужок із нульовою або від’ємною кількістю пар.

Додайте ще одну `if` інструкцію, щоб перевірити, чи `pairs` менше за `1`. Якщо ця умова істинна, поверніть рядок `The number of pairs should be at least 1`.

# --hints--

У вашій `if` функції має бути друга `gen_parentheses` інструкція.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[1]
`) })
```

Друга `if` інструкція має перевіряти, чи `pairs` менше за `1`.

```js
({ test: () => runPython(`
cond = _Node(_code).find_function("gen_parentheses").find_ifs()[1].find_conditions()[0]
assert cond.is_equivalent("pairs < 1") or cond.is_equivalent("1 > pairs")
`) })
```

Потрібно повертати рядок `The number of pairs should be at least 1`, коли `pairs` менше за `1`.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be at least 1'
assert gen_parentheses(0) == msg and gen_parentheses(-1) == msg
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    
    return []
--fcc-editable-region--
```
