---
id: 68da6138d70960c197ec5ef0
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Тепер ви реалізуєте головний цикл BFS. Створіть `while` цикл, який триватиме, доки `queue` не порожній (поки `queue` оцінюється як `True`).

Всередині циклу виведіть `queue`.

# --hints--

У вашій `while` функції має бути `gen_parentheses` цикл.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0]
`) })
```

У вашому `while` циклі має бути умова `queue`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_conditions()[0].is_equivalent("queue")
`) })
```

Ви маєте вивести `queue` всередині вашого `while` циклу.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_call("print(queue)")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    queue = [('', 0, 0)]
    result = []
    
    
    return result
--fcc-editable-region--
```
