---
id: 68da6138d70960c197ec5eec
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Перед тим як реалізувати основний алгоритм, потрібно перевірити вхідні дані. Параметр `pairs` має бути цілим числом, оскільки він представляє кількість пар круглих дужок, які потрібно згенерувати.

Додайте інструкцію `if` на початку своєї функції, щоб перевірити, чи `pairs` не є цілим числом. Для цього використовуйте функцію `isinstance()`.

Якщо умова істинна, поверніть рядок `The number of pairs should be an integer`.

# --hints--

У вашій функції `if` має бути інструкція `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0]
`) })
```

Ваша інструкція `if` має перевіряти, чи `pairs` не є цілим числом, використовуючи `not isinstance(pairs, int)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(pairs, int)")
`) })
```

Потрібно повернути рядок `The number of pairs should be an integer`, коли вхідні дані не є цілим числом.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be an integer'
args = [1.1, "1", [1, 2], {}]
assert  all(gen_parentheses(arg) == msg for arg in args)
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    
    return []
--fcc-editable-region--
```
