---
id: 68da6138d70960c197ec5ef1
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Тепер потрібно визначити, коли ви побудували повну комбінацію круглих дужок. Повна комбінація має довжину, що дорівнює подвоєній кількості пар (оскільки кожна пара складається з однієї відкриваючої та однієї закриваючої дужки).

Всередині вашого `while` циклу додайте інструкцію `if`, щоб перевірити, чи `len(current) == 2 * pairs`. Коли ця умова істинна, додайте `current` до списку `result`.

# --hints--

Усередині вашого `if` циклу має бути інструкція `while`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0]
`) })
```

Ваша інструкція `if` має перевіряти, чи `len(current) == 2 * pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("len(current) == 2 * pairs")
`) })
```

Ви маєте додати `current` до списку `result` всередині нової інструкції `if`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("result.append(current)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        
--fcc-editable-region--    
    return result

print(gen_parentheses(1))
```
