---
id: 68da604b8b6012a33e2344fb
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому практичному занятті ви реалізуєте функцію, яка генерує всі допустимі комбінації круглих дужок, використовуючи підхід пошуку у ширину (BFS). Наприклад, допустимі комбінації двох пар круглих дужок — це `(())` та `()()`.

Почніть зі створення функції з назвою `gen_parentheses` з одним параметром `pairs`. Поки що функція має повертати порожній список.

# --hints--

Ви повинні визначити функцію з назвою `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("gen_parentheses")
`) })
```

Ваша функція `gen_parentheses` має мати один параметр з назвою `pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_args("pairs")
`) })
```

Ваша функція має повертати порожній список.

```js
({ test: () => runPython(`
foo = _Node(_code).find_function("gen_parentheses")
assert foo.has_return("[]") or foo.has_return("list()")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
