---
id: 68e00f88cf82083df8e0497b
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Перш ніж продовжувати роботу над логікою BFS, зробіть виклик `gen_parentheses(1)` і виведіть результат у консоль.

# --hints--

Ви маєте вивести `gen_parentheses(1)` у консоль.

```js
({ test: () => runPython(`
assert _Node(_code).has_call('print(gen_parentheses(1))')
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []

    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)

    return result
--fcc-editable-region--

--fcc-editable-region--
```
