---
id: 68e00c5349cbddfef6dbc053
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Усередині вашого `while` циклу використайте `queue.pop(0)`, щоб видалити та отримати перший елемент із черги. Це реалізує поведінку «перший прийшов — перший пішов» (FIFO), характерну для BFS.

Розпакуйте цей кортеж у три змінні: `current`, `opens_used` і `closes_used`.

# --hints--

Ви повинні використати `queue.pop(0)`, щоб отримати перший елемент із черги.

```js
({ test: () => runPython(`
assert "queue.pop(0)" in str(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0])
`) })
```

Ви повинні розпакувати результат у три змінні: `current`, `opens_used` і `closes_used`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_stmt("current, opens_used, closes_used = queue.pop(0)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        
--fcc-editable-region--
    return result
```
