---
id: 68e45a2b4ab9b2f48b1cc6df
title: Спроектувати сторінку макету цінових планів
challengeType: 25
dashedName: design-a-pricing-plans-layout-page
demoType: onClick
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача**

1. На вашій сторінці має бути елемент `h1` з текстом `Pricing Plans`.
2. На вашій сторінці має бути елемент `div` з класом `pricing-container` під елементом `h1`.
    - Селектор `.pricing-container` має мати властивість `display` зі значенням `flex` і властивість `flex-wrap` зі значенням `wrap`.
3. Всередині елемента `.pricing-container` має бути три елементи `div` з класом `pricing-card` для представлення цінових планів.
    - Один із елементів `.pricing-card` має мати клас `basic-plan` на додаток до класу `pricing-card`.
    - Один із елементів `.pricing-card` має мати клас `pro-plan` на додаток до класу `pricing-card`.
    - Один із елементів `.pricing-card` має мати клас `premium-plan` на додаток до класу `pricing-card`.
4. Ваш елемент `.basic-plan` має містити елемент `h2` з текстом `Basic`.
5. Ваш елемент `.basic-plan` має містити елемент `p` з текстом `$9/month`.
6. Ваш елемент `.pro-plan` має містити елемент `h2` з текстом `Pro`.
7. Ваш елемент `.pro-plan` має містити елемент `p` з текстом `$19/month`.
8. Ваш елемент `.premium-plan` має містити елемент `h2` з текстом `Premium`.
9. Ваш елемент `.premium-plan` має містити елемент `p` з текстом `$29/month`.
10. Кожен із ваших елементів `.pricing-card` має:
    - Використовувати гнучкий контейнер (Flexbox), щоб розташувати вміст у колонку та розподілити простір між дочірніми елементами за допомогою `space-between`.
    - Встановити властивість `flex` у значення `0 0 200px`, щоб надати сталу ширину і запобігти зміні розміру в макеті.
    - Встановити властивість `border` у значення `2px solid black`, щоб побачити, як різні картки займають простір.
11. Елемент `.basic-plan` має з’явитися першим у макеті. Для цього слід використати властивість `order`.
12. Елемент `.pro-plan` має з’явитися другим у макеті. Для цього слід використати властивість `order` і встановити її властивість `flex-grow` у значення `2`, щоб він займав більше місця, ніж інші плани.
13. Елемент `.premium-plan` має бути останнім у макеті. Для цього слід використати властивість `order`.

**Примітка:** Обов’язково підключіть свій файл стилів у HTML і застосуйте CSS.

# --hints--

У вас має бути елемент `h1`.

```js
assert.exists(document.querySelector('h1'));
```

На вашій сторінці має бути елемент `h1` з текстом `Pricing Plans`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Pricing Plans');
```

У вас має бути елемент `div` під елементом `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

Ваш елемент `div` має мати клас `pricing-container`.

```js
assert.exists(document.querySelector('h1 + div.pricing-container'));
```

У вас має бути селектор `.pricing-container`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-container'));
```

Ваш селектор `pricing-container` має мати властивість `display` зі значенням `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.display, 'flex');
```

Ваш селектор `.pricing-container` має мати властивість `flex-wrap` зі значенням `wrap`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.flexWrap, 'wrap');
```

У вас має бути три елементи `div` з класом `pricing-card` всередині елемента `.pricing-container`.

```js
const divEls = document.querySelectorAll('.pricing-container .pricing-card');
assert.lengthOf(divEls, 3);
```

Один із ваших елементів `.pricing-card` має мати клас `basic-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.basic-plan'));
```

Один із ваших елементів `.pricing-card` має мати клас `premium-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.premium-plan'));
```

Один із ваших елементів `.pricing-card` має мати клас `pro-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.pro-plan'));
```

У вас має бути селектор `.pricing-card`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-card'));
```

Ваш селектор `.pricing-card` має мати властивість `display` зі значенням `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.display, 'flex');
```

Ваш селектор `.pricing-card` має мати властивість `flex-direction` зі значенням `column`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flexDirection, 'column');
```

Ваш селектор `.pricing-card` має мати властивість `justify-content` зі значенням `space-between`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.justifyContent, 'space-between');
```

Ваш селектор `.pricing-card` має мати властивість `flex` зі значенням `0 0 200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flex, '0 0 200px');
```

Ваш селектор `.pricing-card` має мати властивість `border` зі значенням `2px solid black`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.border, '2px solid black');
```

Ваш селектор `.basic-plan` має мати властивість `order` зі значенням `0`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.basic-plan')?.order, '0');
```

Ваш елемент `.basic-plan` має містити елемент `h2` з текстом `Basic`.

```js
const h2 = document.querySelector('.basic-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Basic');
```

Ваш елемент `.basic-plan` має містити елемент `p` з текстом `$9/month`.

```js
const p = document.querySelector('.basic-plan > p');
assert.strictEqual(p?.textContent.trim(), '$9/month');
```

Ваш елемент `.pro-plan` має містити елемент `h2` з текстом `Pro`.

```js
const h2 = document.querySelector('.pro-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Pro');
```

Ваш елемент `.pro-plan` має містити елемент `p` з текстом `$19/month`.

```js
const p = document.querySelector('.pro-plan > p');
assert.strictEqual(p?.textContent.trim(), '$19/month');
```

Ваш елемент `.pro-plan` має властивість `order` зі значенням `1`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.order, '1');
```

Ваш елемент `.pro-plan` має властивість `flex-grow` зі значенням `2`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.flexGrow, '2');
```

Ваш елемент `.premium-plan` має містити елемент `h2` з текстом `Premium`.

```js
const h2 = document.querySelector('.premium-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Premium');
```

Ваш елемент `.premium-plan` має містити елемент `p` з текстом `$29/month`.

```js
const p = document.querySelector('.premium-plan > p');
assert.strictEqual(p?.textContent.trim(), '$29/month');
```

Ваш елемент `premium-plan` має бути останнім елементом у макеті і мати властивість `order` зі значенням `2`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.premium-plan')?.order, '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Pricing Plans Layout Page</title>
  </head>
  <body></body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Pricing Plans Layout</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Pricing Plans</h1>
    <div class="pricing-container">
      <div class="pricing-card pro-plan">
        <h2>Pro</h2>
        <p class="price">$19/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card basic-plan">
        <h2>Basic</h2>
        <p class="price">$9/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="cross">✗</span> Advanced analytics</li>
          <li><span class="cross">✗</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card premium-plan">
        <h2>Premium</h2>
        <p class="price">$29/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="check">✓</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  align-items: center;
  justify-content: center;
  flex-direction: column;
  min-height: 100vh;
}

h1 {
  margin-bottom: 20px;
}

h2,
p {
  margin: 0;
}

.pricing-container {
  display: flex;
  flex-wrap: wrap;
  align-content: center;
  border: 2px solid #000;
  gap: 20px;
  width: 80%;
  max-width: 900px;
  height: 70vh;
  padding: 0 20px;
}

.pricing-card {
  background: rgb(247, 247, 247);
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
  padding: 20px;
  flex: 0 0 200px;
  text-align: center;
  transition: transform 0.3s ease;
  display: flex;
  flex-direction: column;
  justify-content: space-between;
  border: 2px solid black;
}

.pricing-card:hover {
  transform: scale(1.05);
}

.price {
  font-size: 1.5rem;
  font-weight: bold;
  color: #3e3e3e;
  margin: 10px;
  border-bottom: 1px solid #433d28;
}

.features {
  list-style-type: none;
  padding: 0;
  text-align: left;
}

.features li {
  padding: 5px 0;
  font-size: 0.95rem;
}

.check {
  color: #219c54;
  font-weight: bold;
  margin-right: 8px;
}

.cross {
  color: crimson;
  font-weight: bold;
  margin-right: 8px;
}

a {
  text-decoration: none;
  display: inline-block;
  margin-top: 15px;
  background: #f1be32;
  color: #000;
  font-weight: bold;
  border: 3px solid #eebe3a;
  padding: 5px;
  transition: background 0.3s;
}

a:hover {
  background: #f7c436;
}

.basic-plan {
  order: 0;
}
.pro-plan {
  order: 1;
  flex-grow: 2;
  background: #f0f8ff;
}
.premium-plan {
  order: 2;
}

@media screen and (max-width: 1000px) {
  .pricing-container {
    height: auto;
    padding-bottom: 20px;
  }

  h1 {
    display: block;
    font-size: 1.5rem;
    text-align: center;
  }

  .pricing-card {
    flex: 0 0 175px;
  }
}

@media screen and (max-width: 650px) {
  .pricing-container {
    flex-direction: column;
    height: auto;
  }

  .pricing-card {
    flex: 1 1 auto;
    width: 100%;
  }
}
```
