---
id: 587d78aa367417b2b2512aed
title: إعلان نوع المستند لوثيقة HTML
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cra98AJ'
forumTopicId: 301095
dashedName: declare-the-doctype-of-an-html-document
---

# --description--

حتى الآن، غطت التحديات عناصر HTML محددة واستخداماتها. مع ذلك، هناك بعض العناصر التي تعطي الهيكل العام لصفحتك، ويجب تضمينها في كل وثيقة HTML.

في أعلى وثيقتك، يجب أن تخبر المتصفح بإصدار HTML الذي تستخدمه صفحتك. HTML لغة تتطور باستمرار وتُحدَّث بانتظام. معظم المتصفحات الرئيسية تدعم أحدث المواصفات، وهي HTML5. مع ذلك، قد تستخدم صفحات الويب القديمة إصدارات سابقة من اللغة.

تخبر المتصفح بهذه المعلومات عن طريق إضافة الوسم `<!DOCTYPE ...>` في السطر الأول، حيث تمثل `...` إصدار HTML. لإصدار HTML5، تستخدم `<!DOCTYPE html>`.

العلامة `!` وكتابة `DOCTYPE` بأحرف كبيرة مهمة، خصوصًا للمتصفحات القديمة. أما `html` فلا تفرق بين الحروف الكبيرة والصغيرة.

بعد ذلك، يجب أن تُغلف بقية شفرة HTML الخاصة بك بوسوم `html`. وسم الفتح `<html>` يوضع مباشرة تحت سطر `<!DOCTYPE html>`، ووسم الإغلاق `</html>` يوضع في نهاية الصفحة.

إليك مثال على هيكل الصفحة. ستضع شفرة HTML الخاصة بك في المساحة بين وسمَي `html`.

```html
<!DOCTYPE html>
<html>

</html>
```

# --instructions--

أضف وسم `DOCTYPE` لإصدار HTML5 في أعلى وثيقة HTML الفارغة في محرر الشفرة. تحته، أضف وسمَي فتح وإغلاق `html`، يحيطان بعنصر `h1`. يمكن أن يحتوي العنوان على أي نص.

# --hints--

يجب أن تتضمن شفرتك وسم `<!DOCTYPE html>`.

```js
assert.match(code,/<!DOCTYPE\s+?html\s*?>/gi);
```

يجب أن يكون هناك عنصر `html` واحد.

```js
assert.lengthOf(document.querySelectorAll('html'), 1);
```

يجب أن تحيط وسوم `html` بعنصر `h1` واحد.

```js
assert.match(code,/<html>\s*?<h1>\s*?.*?\s*?<\/h1>\s*?<\/html>/gi);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <h1> Hello world </h1>
</html>
```
