---
id: 5c6c06847491271903d37cfd
title: استخدام الخاصية value مع أزرار الاختيار ومربعات الاختيار
challengeType: 0
forumTopicId: 301099
dashedName: use-the-value-attribute-with-radio-buttons-and-checkboxes
---

# --description--

عند إرسال نموذج، تُرسل البيانات إلى الخادم وتشمل مدخلات الخيارات المحددة. تُبلغ المدخلات من نوع `radio` و`checkbox` عن قيمها من الخاصية `value`.

على سبيل المثال:

```html
<label for="indoor">
  <input id="indoor" value="indoor" type="radio" name="indoor-outdoor">Indoor
</label>
<label for="outdoor">
  <input id="outdoor" value="outdoor" type="radio" name="indoor-outdoor">Outdoor
</label>
```

هنا، لديك مدخلان من نوع `radio`. عندما يرسل المستخدم النموذج مع اختيار خيار `indoor`، ستتضمن بيانات النموذج السطر: `indoor-outdoor=indoor`. هذا ناتج عن خاصيتي `name` و`value` للمدخل "indoor".

إذا حذفت الخاصية `value`، تستخدم بيانات النموذج المرسلة القيمة الافتراضية وهي `on`. في هذه الحالة، إذا نقر المستخدم خيار "indoor" وأرسل النموذج، ستكون بيانات النموذج الناتجة `indoor-outdoor=on`، وهذا غير مفيد. لذلك يجب تعيين الخاصية `value` إلى شيء يحدد الخيار.

# --instructions--

أضف الخاصية `value` لكل مدخلات `radio` و`checkbox` الموجودة. لا تنشئ أي عناصر جديدة من نوع radio أو checkbox. استخدم نص المسمى الخاص بالمدخل، بأحرف صغيرة، كقيمة للخاصية.

# --hints--

يجب أن يحتوي أحد أزرار الاختيار على الخاصية `value` بالقيمة `indoor`.

```js
const indoorRadioButton = [...document.querySelectorAll('label > input[type="radio"]')].filter(x => x.value === "indoor");
assert.notEmpty(indoorRadioButton)
```

يجب أن يحتوي أحد أزرار الاختيار على الخاصية `value` بالقيمة `outdoor`.

```js
const outdoorRadioButton = [...document.querySelectorAll('label > input[type="radio"]')].filter(x => x.value === "outdoor");
assert.notEmpty(outdoorRadioButton);
```

يجب أن يحتوي أحد مربعات الاختيار على الخاصية `value` بالقيمة `loving`.

```js
const lovingCheckbox = [...document.querySelectorAll('label > input[type="checkbox"]')].filter(x => x.value === "loving");
assert.notEmpty(lovingCheckbox); 
```

يجب أن يحتوي أحد مربعات الاختيار على الخاصية `value` بالقيمة `lazy`.

```js
const lazyCheckbox = [...document.querySelectorAll('label > input[type="checkbox"]')].filter(x => x.value === "lazy");
assert.notEmpty(lazyCheckbox); 
```

يجب أن يحتوي أحد مربعات الاختيار على الخاصية `value` بالقيمة `energetic`.

```js
const energeticCheckbox =  [...document.querySelectorAll('label > input[type="checkbox"]')].filter(x => x.value === "energetic");
assert.notEmpty(energeticCheckbox);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label for="loving"><input id="loving" type="checkbox" name="personality"> Loving</label>
    <label for="lazy"><input id="lazy" type="checkbox" name="personality"> Lazy</label>
    <label for="energetic"><input id="energetic" type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label><br>
    <label for="loving"><input id="loving" type="checkbox" name="personality" value="loving"> Loving</label>
    <label for="lazy"><input id="lazy" type="checkbox" name="personality" value="lazy"> Lazy</label>
    <label for="energetic"><input id="energetic" type="checkbox" name="personality" value="energetic"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
