---
id: bad87fee1348bd9aedf08817
title: إنشاء روابط معطلة باستخدام رمز الهاش
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cMdkytL'
forumTopicId: 18230
dashedName: make-dead-links-using-the-hash-symbol
---

# --description--

أحيانًا تريد إضافة عناصر `a` إلى موقعك قبل أن تعرف إلى أين ستشير الروابط.

هذا مفيد أيضًا عندما تغير سلوك الرابط باستخدام `JavaScript`، وهو ما سنتعلمه لاحقًا.

# --instructions--

القيمة الحالية لخاصية `href` هي رابط يشير إلى "`https://www.freecatphotoapp.com`". استبدل قيمة خاصية `href` بـ `#`، المعروف أيضًا برمز الهاش، لإنشاء رابط معطل.

على سبيل المثال: `href="#"`

# --hints--

يجب أن يكون عنصر `a` الخاص بك رابطًا معطلاً بقيمة خاصية `href` مضبوطة على "#".

```js
assert.strictEqual(document.querySelector('a').getAttribute('href'),'#');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>.</p>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#" target="_blank">cat photos</a>.</p>
  
  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
  
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
