---
id: bad87fee1348bd9aedf08827
title: إنشاء قائمة غير مرتبة بنقاط
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cDKVPuv'
forumTopicId: 16814
dashedName: create-a-bulleted-unordered-list
---

# --description--

يحتوي HTML على عنصر خاص لإنشاء <dfn>قوائم غير مرتبة</dfn>، أو قوائم بنمط نقاط التعداد.

تبدأ القوائم غير المرتبة بعنصر افتتاحي `<ul>`، يليه أي عدد من عناصر `<li>`. وأخيرًا، تُغلق القوائم غير المرتبة بـ `</ul>`.

على سبيل المثال:

```html
<ul>
  <li>milk</li>
  <li>cheese</li>
</ul>
```

سيُنشئ قائمة بنمط نقاط التعداد تحتوي على `milk` و `cheese`.

# --instructions--

احذف آخر عنصرين `p` وأنشئ قائمة غير مرتبة بثلاثة أشياء تحبها القطط في أسفل الصفحة.

# --hints--

أنشئ عنصر `ul`.

```js
assert.isNotEmpty(document.querySelectorAll('ul'));
```

يجب أن يحتوي عنصر `ul` على ثلاثة عناصر `li` بداخله.

```js
assert.lengthOf(document.querySelectorAll('ul li'),3);
```

يجب أن يحتوي عنصر `ul` على وسم الإغلاق.

```js
assert.match(code,/<\/ul>/gi);
assert.match(code,/<ul/gi);
assert.strictEqual(code.match(/<\/ul>/gi).length,code.match(/<ul/gi).length);
```

يجب أن تحتوي عناصر `li` على وسوم إغلاق.

```js
assert.match(code,/<\/li>/gi);
assert.match(code,/<li[\s>]/gi);
assert.strictEqual(code.match(/<\/li>/gi).length,code.match(/<li[\s>]/gi).length);
```

يجب ألا تحتوي عناصر `li` على سلسلة نصية فارغة أو على فراغات فقط.

```js
assert.isEmpty([...document.querySelectorAll('ul li')].filter((item) => item.textContent.trim() === ""));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <ul>
    <li>milk</li>
    <li>mice</li>
    <li>catnip</li>
  </ul>
</main>
```
