---
id: bad87fee1348bd9aedf08834
title: إنشاء مجموعة من أزرار الاختيار
challengeType: 0
forumTopicId: 16822
dashedName: create-a-set-of-radio-buttons
---

# --description--

يمكنك استخدام <dfn>أزرار الاختيار</dfn> للأسئلة التي تريد من المستخدم أن يعطيك فيها إجابة واحدة فقط من بين عدة خيارات.

أزرار الاختيار هي نوع من `input`.

يمكنك تضمين كل زر اختيار داخل عنصر `label` خاص به. عند تغليف عنصر `input` داخل عنصر `label`، سيرتبط مُدخل زر الاختيار تلقائيًا بعنصر المسمى المحيط به.

يجب أن تحمل جميع أزرار الاختيار المرتبطة نفس الخاصية `name` لإنشاء مجموعة أزرار اختيار. بإنشاء مجموعة، يؤدي اختيار أي زر اختيار إلى إلغاء تحديد الأزرار الأخرى في نفس المجموعة تلقائيًا، مما يضمن تقديم إجابة واحدة فقط من المستخدم.

إليك مثالًا على زر اختيار:

```html
<label> 
  <input type="radio" name="indoor-outdoor">Indoor 
</label>
```

تُعتبر أفضل الممارسات تعيين خاصية `for` على عنصر `label`، بقيمة تطابق قيمة الخاصية `id` لعنصر `input`. هذا يسمح لتقنيات المساعدة بإنشاء علاقة مرتبطة بين المسمى وعنصر `input` المرتبط به. على سبيل المثال:

```html
<input id="indoor" type="radio" name="indoor-outdoor">
<label for="indoor">Indoor</label>
```

يمكننا أيضًا تضمين عنصر `input` داخل وسم `label`:

```html
<label for="indoor"> 
  <input id="indoor" type="radio" name="indoor-outdoor">Indoor 
</label>
```

# --instructions--

أضف زوجًا من أزرار الاختيار إلى النموذج الخاص بك، كل واحد مسكن في عنصر `label` خاص به. يجب أن يحتوي أحدهما على الخيار `indoor` والآخر على الخيار `outdoor`. يجب أن يشتركا في الخاصية `name` بقيمة `indoor-outdoor` لإنشاء مجموعة أزرار اختيار.

# --hints--

يجب أن تحتوي صفحتك على عنصرين من أزرار `radio`.

```js
assert.lengthOf(document.querySelectorAll('input[type="radio"]'),2);
```

يجب أن تُعطى أزرار الاختيار الخاصية `name` بقيمة `indoor-outdoor`.

```js
assert.lengthOf([...document.querySelectorAll('input[type="radio"]')].filter(input => input.name === "indoor-outdoor"),2);
```

يجب أن يكون كل من عنصرَي زر الاختيار مسكنًا داخل عنصر `label` خاص به.

```js
assert.lengthOf(document.querySelectorAll('label > input[type="radio"]:only-child'),2);
```

يجب أن يحتوي كل عنصر `label` على وسم إغلاق.

```js
assert.match(code,/<\/label>/g);
assert.match(code,/<label/g);
assert.strictEqual(code.match(/<\/label>/g).length,code.match(/<label/g).length);
```

يجب أن يكون لأحد أزرار الاختيار المسمى `indoor`.

```js
const labelTexts = [...document.querySelectorAll('label')].map(label => label.textContent);
const hasMatch = labelTexts.some(label =>  label.match(/indoor/gi));
assert.isTrue(hasMatch); 
```

يجب أن يكون لأحد أزرار الاختيار المسمى `outdoor`.

```js
const labelTexts = [...document.querySelectorAll('label')].map(label => label.textContent);
const hasMatch = labelTexts.some(label => label.match(/outdoor/gi));
assert.isTrue(hasMatch); 
```

يجب إضافة كل عنصر من أزرار الاختيار داخل وسم `form`.

```js
const labelParent = document.querySelector('label').parentNode;
assert.strictEqual(labelParent.tagName,'FORM');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
   <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
