---
id: bad87fee1348bd9aedf08801
title: الإعلام باستخدام عنصر الفقرة
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/ceZ7DtN'
forumTopicId: 18202
dashedName: inform-with-the-paragraph-element
---

# --description--

عنصر `p` هو العنصر المفضل لنص الفقرة على مواقع الويب. `p` اختصار لكلمة "paragraph".

يمكنك إنشاء عنصر فقرة بهذا الشكل:

```html
<p>I'm a p tag!</p>
```

# --instructions--

أنشئ عنصر `p` أسفل عنصر `h2`، وأضف له النص `Hello Paragraph`.

**ملاحظة:** كاصطلاح، تُكتب جميع وسوم HTML بأحرف صغيرة، مثل `<p></p>` وليس `<P></P>`.

# --hints--

يجب أن يحتوي شفرتك على عنصر `p` صالح.

```js
assert.lengthOf(document.querySelectorAll('p'),1);
```

يجب أن يحتوي عنصر `p` على النص `Hello Paragraph`.

```js
assert.match(document.querySelector('p').textContent,/hello(\s)+paragraph/gi);
```

يجب أن يحتوي عنصر `p` على وسم الإغلاق.

```js
assert.match(code,/<\/p>/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>
<h2>CatPhotoApp</h2>
```

# --solutions--

```html
<h1>Hello World</h1>
<h2>CatPhotoApp</h2>
<p>Hello Paragraph</p>
```
