---
id: bad88fee1348bd9aedf08816
title: الربط بالأقسام الداخلية لصفحة باستخدام عناصر وسم الرابط
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cyrDRUL'
forumTopicId: 301098
dashedName: link-to-internal-sections-of-a-page-with-anchor-elements
---

# --description--

يمكن أيضًا استخدام عناصر `a` (*وسم الرابط*) لإنشاء روابط داخلية للقفز إلى أقسام مختلفة داخل صفحة الويب.

لإنشاء رابط داخلي، عيّن خاصية `href` للرابط إلى رمز التجزئة `#` متبوعًا بقيمة خاصية `id` للعنصر الذي تريد الربط الداخلي به، عادةً في مكان أبعد في الصفحة. ثم أضف نفس خاصية `id` إلى العنصر الذي تربط إليه. الخاصية `id` تصف العنصر بشكل فريد.

فيما يلي مثال على رابط داخلي باستخدام وسم الرابط والعنصر الهدف:

```html
<a href="#contacts-header">Contacts</a>
...
<h2 id="contacts-header">Contacts</h2>
```

عندما ينقر المستخدمون على رابط `Contacts`، سيتم نقلهم إلى قسم صفحة الويب الذي يحتوي على عنصر العنوان **Contacts**.

# --instructions--

غيّر الرابط الخارجي إلى رابط داخلي بتغيير خاصية `href` إلى `#footer` والنص من `cat photos` إلى `Jump to Bottom`.

أزل الخاصية `target="_blank"` من وسم الرابط لأن ذلك يفتح المستند المرتبط في تبويب نافذة جديدة.

ثم أضف خاصية `id` بقيمة `footer` إلى عنصر `<footer>` في أسفل الصفحة.

# --hints--

يجب أن يكون هناك وسم رابط واحد فقط في صفحتك.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

يجب أن يكون هناك وسم `footer` واحد فقط في صفحتك.

```js
assert.lengthOf(document.querySelectorAll('footer'),1);
```

يجب أن يحتوي وسم `a` على خاصية `href` مضبوطة على "#footer".

```js
assert.strictEqual(document.querySelector('a').getAttribute('href'), '#footer');
```

يجب ألا يحتوي وسم `a` على خاصية `target`.

```js
assert.notExists(document.querySelector('a').getAttribute('target'));
```

يجب أن يكون نص وسم `a` هو "Jump to Bottom".

```js
assert.match(document.querySelector('a').textContent,/Jump to Bottom/gi);
```

يجب أن يحتوي وسم `footer` على خاصية `id` مضبوطة على "footer".

```js
assert.strictEqual(document.querySelector('footer').getAttribute('id'),'footer');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
  <p>Meowwww loved it, hated it, loved it, hated it yet spill litter box, scratch at owner, destroy all furniture, especially couch or lay on arms while you're using the keyboard. Missing until dinner time toy mouse squeak roll over. With tail in the air lounge in doorway. Man running from cops stops to pet cats, goes to jail.</p>
  <p>Intently stare at the same spot poop in the plant pot but kitten is playing with dead mouse. Get video posted to internet for chasing red dot leave fur on owners clothes meow to be let out and mesmerizing birds leave fur on owners clothes or favor packaging over toy so purr for no reason. Meow to be let out play time intently sniff hand run outside as soon as door open yet destroy couch.</p>

</main>

<footer>Copyright Cat Photo App</footer>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="#footer">Jump to Bottom</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
  <p>Meowwww loved it, hated it, loved it, hated it yet spill litter box, scratch at owner, destroy all furniture, especially couch or lay on arms while you're using the keyboard. Missing until dinner time toy mouse squeak roll over. With tail in the air lounge in doorway. Man running from cops stops to pet cats, goes to jail.</p>
  <p>Intently stare at the same spot poop in the plant pot but kitten is playing with dead mouse. Get video posted to internet for chasing red dot leave fur on owners clothes meow to be let out and mesmerizing birds leave fur on owners clothes or favor packaging over toy so purr for no reason. Meow to be let out play time intently sniff hand run outside as soon as door open yet destroy couch.</p>

</main>

<footer id="footer">Copyright Cat Photo App</footer>
```
