---
id: bad87fee1348bd9aedf08812
title: إضافة صور إلى موقعك
challengeType: 0
forumTopicId: 16640
dashedName: add-images-to-your-website
---

# --description--

يمكنك إضافة صور إلى موقعك باستخدام عنصر `img`، وتحديد عنوان URL لصورة معينة باستخدام الخاصية `src`.

مثال على ذلك هو:

```html
<img src="https://www.freecatphotoapp.com/your-image.jpg">
```

لاحظ أن `img` عنصر فارغ.

يجب أن تحتوي جميع عناصر `img` على خاصية `alt`. النص داخل خاصية `alt` يُستخدم لقارئ الشاشة لتحسين إمكانية الوصول ويُعرض إذا فشلت الصورة في التحميل.

**ملاحظة:** إذا كانت الصورة للزينة فقط، فإن استخدام خاصية `alt` فارغة يعد من أفضل الممارسات.

من الأفضل ألا تحتوي خاصية `alt` على أحرف خاصة إلا إذا كان ذلك ضروريًا.

لنُضف خاصية `alt` إلى مثال `img` أعلاه:

```html
<img src="https://www.freecatphotoapp.com/your-image.jpg" alt="freeCodeCamp logo">
```

# --instructions--

لنحاول إضافة صورة إلى موقعنا:

داخل عنصر `main` الموجود، أدخل عنصر `img` قبل عناصر `p` الموجودة.

الآن عرّف خاصية `src` بحيث تشير إلى العنوان `https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

أخيرًا، لا تنسَ إعطاء عنصر `img` خاصية `alt` بنص مناسب.

# --hints--

يجب أن تحتوي صفحتك على عنصر صورة.

```js
assert.exists(document.querySelector('img'));
```

يجب أن تحتوي صورتك على خاصية `src` تشير إلى صورة القط الصغير.

```js
const url = document.querySelector('img').getAttribute('src');
assert.match(url,/^https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg$/i);
```

يجب ألا تكون خاصية `alt` في عنصر الصورة فارغة.

```js
assert.exists(document.querySelector('img').getAttribute('alt'));
assert.isNotEmpty(document.querySelector('img').getAttribute('alt'));
assert.match(__helpers.removeWhiteSpace(code),/<(?:img|IMG)\S*alt=(['"])(?!\1|>)\S+\1\S*\/?>/)
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>


  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
