---
id: bad87fee1348bd9aede08817
title: التسكين عنصر وسم الرابط داخل فقرة
challengeType: 0
forumTopicId: 18244
dashedName: nest-an-anchor-element-within-a-paragraph
---

# --description--

يمكنك تسكين الروابط داخل عناصر نصية أخرى.

```html
<p>
  Here's a <a target="_blank" href="https://www.freecodecamp.org"> link to www.freecodecamp.org</a> for you to follow.
</p>
```

لنحلل المثال. النص العادي ملفوف داخل عنصر `p`:

```html
<p> Here's a ... for you to follow. </p>
```

التالي هو عنصر *وسم الرابط* `<a>` (الذي يتطلب وسم إغلاق `</a>`):

```html
<a> ... </a>
```

`target` خاصية في وسم الرابط تحدد مكان فتح الرابط. القيمة `_blank` تحدد فتح الرابط في تبويب جديد. الخاصية `href` هي خاصية في وسم الرابط تحتوي على عنوان URL للرابط:

```html
<a href="https://www.freecodecamp.org" target="_blank"> ... </a>
```

النص `link to www.freecodecamp.org` داخل عنصر `a` يسمى <dfn>نص وسم الرابط</dfn>، وسيعرض الرابط للنقر عليه:

```html
<a href=" ... " target="...">link to freecodecamp.org</a>
```

الناتج النهائي للمثال سيبدو هكذا:

إليك <a href="https://www.freecodecamp.org" target="_blank">رابط إلى www.freecodecamp.org</a> لتتابعه.

# --instructions--

سكن عنصر `a` الموجود داخل عنصر `p` جديد. لا تنشئ وسم رابط جديد. يجب أن تحتوي الفقرة الجديدة على نص يقول `View more cat photos`، حيث يكون `cat photos` رابطًا، والباقي نص عادي.

# --hints--

يجب أن يكون لديك عنصر `a` واحد فقط.

```js
assert.lengthOf(document.querySelectorAll('a'), 1 );
```

يجب أن يربط عنصر `a` بـ "`https://www.freecatphotoapp.com`".

```js
assert.lengthOf(document.querySelectorAll('a[href="https://www.freecatphotoapp.com"]'),1);
```

يجب أن يحتوي عنصر `a` على نص وسم الرابط `cat photos`

```js
assert.match(document.querySelector('a').textContent,/cat\sphotos/gi)
```

يجب أن تنشئ عنصر `p` جديد. يجب أن يكون هناك على الأقل 3 وسوم `p` في شفرة HTML الخاصة بك.

```js
assert.lengthOf(document.querySelectorAll('p'),3)
```

يجب أن يكون عنصر `a` مسكنًا داخل عنصر `p` الجديد.

```js
const anchorParent = document.querySelector('a[href="https://www.freecatphotoapp.com"]').parentNode;
assert.strictEqual(anchorParent.tagName,"P")
```

يجب أن يحتوي عنصر `p` على النص `View more ` (مع وجود مسافة بعده).

```js
const textContent = document.querySelector('a[href="https://www.freecatphotoapp.com"]').parentNode.textContent;
assert.match(textContent,/View\smore\s/gi);
```

يجب ألا يحتوي عنصر `a` على النص `View more` <em>أبدًا</em>.

```js
assert.notMatch(document.querySelector('a').textContent,/View\smore/gi); 
```

يجب أن يحتوي كل عنصر `p` على وسم إغلاق.

```js
assert.match(code,/<\/p>/g);
assert.match(code,/<p/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

يجب أن يحتوي كل عنصر `a` على وسم إغلاق.

```js
assert.match(code,/<\/a>/g);
assert.match(code,/<a/g);
assert.strictEqual(code.match(/<\/a>/g).length,code.match(/<a/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>View more <a target="_blank" href="https://www.freecatphotoapp.com">cat photos</a></p>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
