---
id: bad87fee1348bd9aedf08816
title: الربط بصفحات خارجية باستخدام عناصر وسم الرابط
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/c8EkncB'
forumTopicId: 18226
dashedName: link-to-external-pages-with-anchor-elements
---

# --description--

يمكنك استخدام عناصر `a` (*وسم الرابط*) للربط بمحتوى خارج صفحة الويب الخاصة بك.

تحتاج عناصر `a` إلى عنوان ويب وجهة يُسمى خاصية `href`. كما تحتاج إلى نص وسم الرابط. إليك مثالًا:

```html
<a href="https://www.freecodecamp.org">this links to freecodecamp.org</a>
```

بعدها سيعرض متصفحك النص `this links to freecodecamp.org` كرابط يمكنك النقر عليه. وهذا الرابط سيأخذك إلى عنوان الويب `https://www.freecodecamp.org`.

# --instructions--

أنشئ عنصر `a` يربط بـ `https://www.freecatphotoapp.com` ويحتوي على "cat photos" كنص وسم الرابط.

# --hints--

يجب أن يحتوي عنصر `a` على نص وسم الرابط `cat photos`.

```js
assert.match(document.querySelector('a').textContent,/cat photos/gi);
```

تحتاج إلى عنصر `a` يربط بـ `https://www.freecatphotoapp.com`

```js
assert.match(document.querySelector('a').getAttribute('href'),/^https?:\/\/(www\.)?freecatphotoapp\.com\/?$/i);
```

يجب أن يحتوي عنصر `a` على وسم الإغلاق.

```js
assert.match(code,/<\/a>/g);
assert.strictEqual(code.match(/<\/a>/g).length,code.match(/<a/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>



  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  
  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
  
  <a href="https://www.freecatphotoapp.com">cat photos</a>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
