---
id: bad87fee1348bd9aedf08802
title: إزالة التعليقات من HTML
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cBmG9T7'
forumTopicId: 18329
dashedName: uncomment-html
---

# --description--

التعليق هو طريقة تترك بها تعليقات لمطوِّرين آخرين داخل شفرتك دون التأثير على الناتج المعروض للمستخدم النهائي.

التعليق أيضًا طريقة مناسبة لجعل الشفرة غير نشطة دون الحاجة إلى حذفها بالكامل.

التعليقات في HTML تبدأ بـ `<!--` وتنتهي بـ `-->`

# --instructions--

أزل التعليقات عن عناصر `h1` و `h2` و `p` الخاصة بك.

# --hints--

يجب أن يكون عنصر `h1` مرئيًا على الصفحة بعد إزالة التعليق عنه.

```js
assert.notEmpty(document.querySelectorAll('h1'));
```

يجب أن يكون عنصر `h2` مرئيًا على الصفحة بعد إزالة التعليق عنه.

```js
assert.notEmpty(document.querySelectorAll('h2'));
```

يجب أن يكون عنصر `p` مرئيًا على الصفحة بعد إزالة التعليق عنه.

```js
assert.notEmpty(document.querySelectorAll('p'));
```

لا يجب أن تظهر وسوم التعليق المتبقية على الصفحة (أي `-->`).

```js
const elements = document.querySelectorAll('*');
const potentialComments = Array.from(elements).filter(el => el.textContent.includes('-->'));
assert.notExists(potentialComments[1])
```

# --seed--

## --seed-contents--

```html
<!--
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
-->
```

# --solutions--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```
