---
id: bad87fed1348bd9aedf08833
title: حذف عناصر HTML
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/ckK73C9'
forumTopicId: 17559
dashedName: delete-html-elements
---

# --description--

هاتفنا لا يحتوي على مساحة عمودية كبيرة.

لنحذف العناصر غير الضرورية لنتمكن من بدء بناء تطبيق CatPhotoApp.

# --instructions--

احذف عنصر `h1` لتبسيط العرض.

# --hints--

يجب حذف عنصر `h1`.

```js
assert.notMatch(code,/<h1>/gi);
assert.notMatch(code,/<\/h1>/gi);
```

يجب أن يكون عنصر `h2` موجودًا في الصفحة.

```js
assert.match(code,/<h2>[\w\W]*<\/h2>/gi);
```

يجب أن يكون عنصر `p` موجودًا في الصفحة.

```js
assert.match(code,/<p>[\w\W]*<\/p>/gi);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```

# --solutions--

```html
<h2>CatPhotoApp</h2><p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```
