---
id: bad87fee1348bd9aedf08820
title: تحويل صورة إلى رابط
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cRdBnUr'
forumTopicId: 18327
dashedName: turn-an-image-into-a-link
---

# --description--

يمكنك تحويل العناصر إلى روابط عن طريق تسكينها داخل عنصر `a`.

سكن صورتك داخل عنصر `a`. إليك مثالًا:

```html
<a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
```

تذكر استخدام `#` كخاصية `href` لعنصر `a` لجعله رابطًا غير فعال.

# --instructions--

ضع عنصر الصورة الموجود داخل عنصر `a` (*وسم الرابط أو وسم الرابط للنص العالي*).

بعد ذلك، مرر مؤشر الفأرة فوق صورتك. يجب أن يتحول مؤشر الفأرة العادي إلى مؤشر النقر على الرابط. أصبحت الصورة الآن رابطًا.

# --hints--

يجب أن يكون عنصر `img` الموجود مسكنًا داخل عنصر `a`.

```js
const anchor = document.querySelectorAll('a')[1];
const children = anchor.querySelectorAll("img");
assert.notEmpty(children);
```

يجب أن يكون عنصر `a` الخاص بك رابطًا غير فعال مع خاصية `href` مضبوطة على `#`.

```js
const anchor = document.querySelectorAll('a')[1];
const parentHREF = anchor.querySelector("img").parentNode.getAttribute('href');
assert.match(parentHREF,new RegExp('#'));
```

يجب أن يحتوي كل عنصر `a` على وسم إغلاق.

```js
assert.match(code,/<\/a>/g);
assert.match(code,/<a/g);
assert.strictEqual(code.match(/<\/a>/g).length, code.match(/<a/g).length)
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
