---
id: bad87fee1348bd9aede08830
title: إنشاء عنصر نموذج
challengeType: 0
forumTopicId: 16817
dashedName: create-a-form-element
---

# --description--

يمكنك بناء نماذج ويب تُرسل البيانات فعليًا إلى خادم باستخدام HTML نقي فقط. يمكنك فعل ذلك بتحديد خاصية `action` على عنصر `form` الخاص بك.

على سبيل المثال:

```html
<form action="url-where-you-want-to-submit-form-data">
  <input>
</form>
```

# --instructions--

قم بتسكين عنصر الـ `input` الموجود داخل عنصر `form` وقم بتعيين `"https://www.freecatphotoapp.com/submit-cat-photo"` إلى خاصية `action` في عنصر الـ `form`.

# --hints--

يجب أن يكون عنصر الـ `input` الموجود مسكنًا داخل عنصر `form`.

```js
const inputElem = document.querySelector('form input');
assert.strictEqual(inputElem.getAttribute('type'),'text'); 
assert.strictEqual(inputElem.getAttribute('placeholder'),'cat photo URL');
```

يجب أن يحتوي نموذجك `form` على خاصية `action` مضبوطة على `https://www.freecatphotoapp.com/submit-cat-photo`.

```js
const action = document.querySelector('form').getAttribute('action');
assert.match(action,/^https:\/\/(www\.)?freecatphotoapp\.com\/submit-cat-photo$/i);
```

يجب أن يحتوي عنصر الـ `form` على وسم افتتاح ووسم إغلاق صحيحين.

```js
assert.match(code,/<\/form>/g);
assert.match(code,/<form [^<]*>/g);
assert.strictEqual(code.match(/<\/form>/g).length,code.match(/<form [^<]*>/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text" placeholder="cat photo URL">
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL">
  </form>
</main>
```
