---
id: bad87fee1348bd9aedc08830
title: استخدام HTML5 لجعل الحقل مطلوبًا
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cMd4EcQ'
forumTopicId: 18360
dashedName: use-html5-to-require-a-field
---

# --description--

يمكنك جعل حقول نموذج معينة مطلوبة بحيث لا يستطيع المستخدم إرسال النموذج حتى يملأها.

على سبيل المثال، إذا أردت جعل حقل إدخال نصي مطلوبًا، يمكنك فقط إضافة الخاصية `required` داخل عنصر `input` الخاص بك، مثل هذا: `<input type="text" required>`

# --instructions--

اجعل حقل `input` النصي خاصتك حقلًا `required`، حتى لا يستطيع المستخدم إرسال النموذج دون إكمال هذا الحقل.

ثم حاول إرسال النموذج دون إدخال أي نص. لاحظ كيف يُنبّهك نموذج HTML5 أن الحقل مطلوب؟

# --hints--

يجب أن يحتوي عنصر `input` النصي على الخاصية `required`.

```js
assert.isTrue(document.querySelector('input').required);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL">
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" required placeholder="cat photo URL">
    <button type="submit">Submit</button>
  </form>
</main>
```
