---
id: bad87fee1348bd9aedf08828
title: إنشاء قائمة مرتبة
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cQ3B8TM'
forumTopicId: 16824
dashedName: create-an-ordered-list
---

# --description--

يحتوي HTML على عنصر خاص آخر لإنشاء <dfn>قوائم مرتبة</dfn>، أو قوائم مرقمة.

تبدأ القوائم المرتبة بعنصر فتح `<ol>`، يليه أي عدد من عناصر `<li>`. وأخيرًا تُغلق القوائم المرتبة بوسم الإغلاق `</ol>`.

على سبيل المثال:

```html
<ol>
  <li>Garfield</li>
  <li>Sylvester</li>
</ol>
```

سيُنشئ قائمة مرقمة تحتوي على `Garfield` و `Sylvester`.

# --instructions--

أنشئ قائمة مرتبة بأهم 3 أشياء تكرهها القطط أكثر شيء.

# --hints--

يجب أن تحتوي على قائمة مرتبة لـ `Top 3 things cats hate:`

```js
const previousElement = document.querySelector('ol').previousElementSibling; 
assert.match(previousElement.textContent,/Top 3 things cats hate:/i);
```

يجب أن تحتوي على قائمة غير مرتبة لـ `Things cats love:`

```js
const previousElement = document.querySelector('ul').previousElementSibling; 
assert.match(previousElement.textContent,/Things cats love:/i);
```

يجب أن تحتوي على عنصر `ul` واحد فقط.

```js
assert.lengthOf(document.querySelectorAll('ul'), 1);
```

يجب أن تحتوي على عنصر `ol` واحد فقط.

```js
assert.lengthOf(document.querySelectorAll('ol'), 1);
```

يجب أن تحتوي على ثلاثة عناصر `li` داخل عنصر `ul` الخاص بك.

```js
assert.lengthOf(document.querySelectorAll('ul li'), 3);
```

يجب أن تحتوي على ثلاثة عناصر `li` داخل عنصر `ol` الخاص بك.

```js
assert.lengthOf(document.querySelectorAll('ol li'), 3);
```

يجب أن يحتوي عنصر `ul` على وسم إغلاق.

```js
assert.match(code,/<\/ul>/g);
assert.strictEqual(code.match(/<\/ul>/g).length ,code.match(/<ul>/g).length);
```

يجب أن يحتوي عنصر `ol` على وسم إغلاق.

```js
assert.match(code,/<\/ol>/g);
assert.strictEqual(code.match(/<\/ol>/g).length ,code.match(/<ol>/g).length);
```

يجب أن يحتوي عنصر `li` على وسم إغلاق.

```js
assert.match(code,/<\/li>/g);
assert.match(code,/<li>/g);
assert.strictEqual(code.match(/<\/li>/g).length ,code.match(/<li>/g).length);
```

يجب ألا تكون عناصر `li` في القائمة غير المرتبة فارغة.

```js
[...document.querySelectorAll('ul li')].forEach((val) =>
  assert.isNotEmpty(__helpers.removeWhiteSpace(val.textContent))
);
```

يجب ألا تكون عناصر `li` في القائمة المرتبة فارغة.

```js
[...document.querySelectorAll('ol li')].forEach((val) =>
  assert.isNotEmpty(__helpers.removeWhiteSpace(val.textContent))
);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>

</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>hate 1</li>
    <li>hate 2</li>
    <li>hate 3</li>
  </ol>
</main>
```
