---
id: bad87fee1348bd9aedf08804
title: إضافة تعليق على HTML
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cGyGbca'
forumTopicId: 16782
dashedName: comment-out-html
---

# --description--

تذكر أنه لبدء تعليق، يجب استخدام `<!--` ولإنهاء التعليق، يجب استخدام `-->`

هنا تحتاج إلى إنهاء التعليق قبل أن يبدأ عنصر `h2` الخاص بك.

# --instructions--

أضف تعليقًا على عنصر `h1` وعنصر `p`، لكن لا تضف تعليقًا على عنصر `h2`.

# --hints--

يجب إضافة تعليق على عنصر `h1` بحيث لا يظهر على الصفحة.

```js
assert.isEmpty(document.querySelectorAll('h1'));
```

يجب ألا يُضاف تعليق على عنصر `h2` لكي يظهر على الصفحة.

```js
assert.isNotEmpty(document.querySelectorAll('h2'));
```

يجب إضافة تعليق على عنصر `p` بحيث لا يظهر على الصفحة.

```js
assert.isEmpty(document.querySelectorAll('p'));
```

يجب إغلاق كل تعليق باستخدام `-->`.

```js
assert.isAbove(code.match(/[^fc]-->/g).length, 1);
```

لا تغير ترتيب عناصر `h1` أو `h2` أو `p` في الشفرة.

```js
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[0],'<h1>');
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[1],'<h2>');
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[2],'<p>');
```

# --seed--

## --seed-contents--

```html
<!--
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
-->
```

# --solutions--

```html
<!--<h1>Hello World</h1>-->
<h2>CatPhotoApp</h2> 
<!--<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p> -->
```
