---
id: bad87fee1348bd9aedf08830
title: إضافة نص عنصر نائب إلى حقل نصي
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cKdJDhg'
forumTopicId: 16647
dashedName: add-placeholder-text-to-a-text-field
---

# --description--

نص العنصر النائب هو ما يُعرض في عنصر `input` قبل أن يُدخل المستخدم أي شيء.

يمكنك إنشاء نص عنصر نائب كما يلي:

```html
<input type="text" placeholder="this is placeholder text">
```

**ملاحظة:** تذكر أن `input` هو عنصر فارغ.

# --instructions--

عيّن قيمة `placeholder` في عنصر `input` النصي إلى "cat photo URL".

# --hints--

يجب إضافة خاصية `placeholder` إلى عنصر `input` النصي الموجود.

```js
assert.notEmpty(document.querySelectorAll('input[placeholder]'));
```

يجب تعيين قيمة خاصية `placeholder` إلى `cat photo URL`.

```js
assert.exists(document.querySelector('input'));
assert.exists(document.querySelector('input').getAttribute('placeholder'));
const placeholder = document.querySelector('input').getAttribute('placeholder');
assert.match(placeholder,/cat\s+photo\s+URL/gi);
```

يجب ألا يحتوي عنصر `input` النهائي على وسم إغلاق.

```js
assert.notMatch(code,/<input.*\/?>.*<\/input>/gi);
```

يجب أن يحتوي عنصر `input` النهائي على صيغة نحوية صحيحة.

```js
assert.notEmpty(document.querySelectorAll('input[type=text]'));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text">
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text" placeholder="cat photo URL">
</main>
```
