---
id: bad87fee1348bd9aedf08835
title: إنشاء مجموعة من مربعات الاختيار
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cqrkJsp'
forumTopicId: 16821
dashedName: create-a-set-of-checkboxes
---

# --description--

تستخدم النماذج عادةً <dfn>مربعات اختيار</dfn> للأسئلة التي قد يكون لها أكثر من إجابة.

مربعات الاختيار هي نوع من `input`.

يمكن أن يكون كل مربع اختيار مسكنًا داخل عنصر `label` خاص به. عند تغليف عنصر `input` داخل عنصر `label`، يرتبط مربع الاختيار تلقائيًا بعنصر المسمى المحيط به.

يجب أن تشترك جميع مربعات الاختيار المرتبطة بنفس الخاصية `name`.

يُعتبر من أفضل الممارسات تحديد العلاقة بين مربع اختيار `input` و`label` المقابل له بشكل صريح عن طريق تعيين خاصية `for` في عنصر `label` لتطابق خاصية `id` في عنصر `input` المرتبط.

إليك مثالًا على مربع اختيار:

```html
<label for="loving"><input id="loving" type="checkbox" name="personality"> Loving</label>
```

# --instructions--

أضف إلى النموذج مجموعة من ثلاثة مربعات اختيار. يجب أن يكون كل مربع اختيار مسكنًا داخل عنصر `label` خاص به. يجب أن تشترك الثلاثة جميعًا في الخاصية `name` بقيمة `personality`.

# --hints--

يجب أن تحتوي صفحتك على ثلاثة عناصر مربع اختيار.

```js
assert.lengthOf(document.querySelectorAll('input[type="checkbox"]'),3);
```

يجب أن يكون كل من عناصر مربعات الاختيار الثلاثة مسكنًا داخل عنصر `label` خاص به.

```js
assert.lengthOf(document.querySelectorAll('label > input[type="checkbox"]:only-child'),3);
```

تأكد من أن لكل عنصر `label` وسم إغلاق.

```js
assert.match(code,/<\/label>/g);
assert.match(code,/<label/g);
assert.strictEqual(code.match(/<\/label>/g).length,code.match(/<label/g).length)
```

يجب أن تُعطى مربعات الاختيار الخاصية `name` بقيمة `personality`.

```js
assert.lengthOf([...document.querySelectorAll('label > input[type="checkbox"]')].filter(x => x.name === "personality"),3);
```

يجب إضافة كل مربع اختيار داخل وسم `form`.

```js
assert.strictEqual(document.querySelector('label').parentNode.tagName,'FORM');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label for="playful"><input id="playful" type="checkbox" name="personality">Playful</label>
    <label for="lazy"><input id="lazy" type="checkbox" 
name="personality">Lazy</label>
    <label for="evil"><input id="evil" type="checkbox" 
name="personality">Evil</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
