---
id: bad87fee1348bd9aecf08801
title: مقدمة إلى عناصر HTML5
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cBkZGpt7'
forumTopicId: 301097
dashedName: introduction-to-html5-elements
---

# --description--

يُقدّم HTML5 وسمًا أكثر وصفًا في HTML. تشمل هذه الوسوم `main`، و`header`، و`footer`، و`nav`، و`video`، و`article`، و`section` وغيرها.

تُعطي هذه الوسوم هيكلًا وصفيًا لشفرة HTML، وتجعل شفرتك أسهل في القراءة، وتساعد في تحسين محركات البحث (SEO) وإمكانية الوصول. وسم `main` في HTML5 يساعد محركات البحث والمطورين الآخرين في العثور على المحتوى الرئيسي لصفحتك.

مثال على الاستخدام، عنصر `main` يحتوي على عنصرين فرعيين مسكنين بداخله:

```html
<main> 
  <h1>Hello World</h1>
  <p>Hello Paragraph</p>
</main>
```

**ملاحظة:** العديد من الوسوم الجديدة في HTML5 وفوائدها مغطاة في قسم إمكانية الوصول التطبيقي.

# --instructions--

أنشئ عنصر `p` ثانيًا يحتوي على نص kitty ipsum التالي: `Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.`

بعد ذلك، أنشئ عنصر `main` وسكّن فقط عنصري `p` داخل عنصر `main`.

# --hints--

يجب أن يكون لديك عنصران `p` يحتويان على نص Kitty Ipsum.

```js
assert.lengthOf(document.querySelectorAll('p'),2);
```

يجب أن يحتوي كل عنصر `p` على وسم إغلاق.

```js
assert.match(code,/<\/p>/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

يجب أن يحتوي عنصر `p` على الكلمات القليلة الأولى من نص `kitty ipsum` الإضافي المقدم.

```js
assert.match(document.querySelectorAll('p')[1].textContent,/Purr\s+jump\s+eat/gi);
```

يجب أن يحتوي شفرتك على عنصر `main` واحد.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

يجب أن يحتوي عنصر `main` على عنصرين فقرتين كأبناء.

```js
const main = document.querySelector('main');
const children = main.querySelectorAll("p"); 
assert.lengthOf(children,2);
```

يجب أن يأتي وسم الفتح `main` قبل وسم الفقرة الأولى.

```js
assert.match(code,(/<main>\s*?<p>/g));
```

يجب أن يأتي وسم الإغلاق `main` بعد وسم إغلاق الفقرة الثانية.

```js
assert.match(code,(/<\/p>\s*?<\/main>/g));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
