---
id: bad87fee1348bd9aedf08833
title: املأ الفراغ بنص العنصر النائب
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cgR7Dc7'
forumTopicId: 18178
dashedName: fill-in-the-blank-with-placeholder-text
---

# --description--

يستخدم مطوّرو الويب تقليديًا <dfn>نص lorem ipsum</dfn> كنص عنصر نائب. يُستخرج نص lorem ipsum عشوائيًا من مقطع مشهور لشيشرون من روما القديمة.

استُخدم نص lorem ipsum كنص عنصر نائب من قبل الطابعين منذ القرن السادس عشر، ولا تزال هذه العادة مستمرة على الويب.

خمسة قرون فترة طويلة بما فيه الكفاية. بما أننا نبني تطبيق CatPhotoApp، فلنستخدم شيئًا يسمى نص "kitty ipsum".

# --instructions--

استبدل النص داخل عنصر `p` الخاص بك بالكلمات القليلة الأولى من نص kitty ipsum هذا: `Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.`

# --hints--

يجب أن يحتوي عنصر `p` الخاص بك على الكلمات القليلة الأولى من نص "kitty ipsum" المقدم.

```js
assert.match(document.querySelector('p').textContent,/Kitty(\s)+ipsum/gi);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Hello Paragraph</p>
```

# --solutions--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff</p>
```
