---
id: 587d824b367417b2b2512c49
title: اختبار القيمة الحقيقية
challengeType: 2
forumTopicId: 301596
dashedName: test-for-truthiness
---

# --description--

كتذكير، يُبنى هذا المشروع على مشروع بداية تم نسخه من <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

تختبر الدالة `isTrue()` القيمة المنطقية `true`، وتمرر الدالة `isNotTrue()` عندما تُعطى أي قيمة غير القيمة المنطقية `true`.

```js
assert.isTrue(true, 'This will pass with the boolean value true');
assert.isTrue('true', 'This will NOT pass with the string value "true"');
assert.isTrue(1, 'This will NOT pass with the number value 1');
```

توجد أيضًا الدالتان `isFalse()` و`isNotFalse()`، وتعملان بطريقة مشابهة لنظيرتيهما الحقيقيتين، لكنهما تبحثان عن القيمة المنطقية `false`.

# --instructions--

داخل الملف `tests/1_unit-tests.js` تحت الاختبار الموسوم بـ `#4` في مجموعة `Basic Assertions`، غيّر كل `assert` إلى إما `assert.isTrue` أو `assert.isNotTrue` لجعل الاختبار ينجح (يجب أن يُقيّم إلى `true`). لا تغير المعلمات الممررة إلى عمليات التحقق.

# --hints--

يجب أن تنجح جميع الاختبارات.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

اختر الطريقة الصحيحة للتحقق الأول - `isTrue` أم `isNotTrue`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isTrue', 'True is true');
```

اختر الطريقة الصحيحة للتحقق الثاني - `isTrue` أم `isNotTrue`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'isTrue',
    'Double negation of a truthy value is true'
  );
```

اختر الطريقة الصحيحة للتحقق الثالث - `isTrue` أم `isNotTrue`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'isNotTrue',
    'A truthy object is not true - neither is a false one'
  );
```

