---
id: 587d824e367417b2b2512c58
title: تشغيل اختبارات وظيفية على نقاط نهاية API باستخدام Chai-HTTP
challengeType: 2
forumTopicId: 301593
dashedName: run-functional-tests-on-api-endpoints-using-chai-http
---

# --description--

كتذكير، يُبنى هذا المشروع على مشروع بداية مستنسخ من <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

تسمح لك Mocha باختبار العمليات غير المتزامنة مثل الاستدعاءات إلى نقاط نهاية API باستخدام إضافة تسمى `chai-http`.

فيما يلي مثال على اختبار يستخدم `chai-http` لمجموعة تسمى `'GET /hello?name=[name] => "hello [name]"'`:

```js
suite('GET /hello?name=[name] => "hello [name]"', function () {
  test('?name=John', function (done) {
    chai
      .request(server)
      .keepOpen()
      .get('/hello?name=John')
      .end(function (err, res) {
        assert.equal(res.status, 200, 'Response status should be 200');
        assert.equal(res.text, 'hello John', 'Response should be "hello John"');
        done();
      });
  });
});
```

يرسل الاختبار طلب `GET` إلى الخادم مع اسم كسلسلة استعلام في عنوان URL (`?name=John`). في دالة الاستدعاء الراجعة لطريقة `end`، يتم استقبال كائن الاستجابة (`res`) الذي يحتوي على خاصية `status`.

الـ `assert.equal` الأول يتحقق مما إذا كانت الحالة تساوي `200`. الـ `assert.equal` الثاني يتحقق من أن سلسلة الاستجابة (`res.text`) تساوي `"hello John"`.

لاحظ أيضًا المعلمة `done` في دالة الاستدعاء الراجعة للاختبار. استدعاؤها بدون معلمة في نهاية الاختبار ضروري للإشارة إلى اكتمال العملية غير المتزامنة.

أخيرًا، لاحظ طريقة `keepOpen` مباشرة بعد طريقة `request`. عادةً ما تشغل اختباراتك من سطر الأوامر، أو كجزء من عملية تكامل آلية، ويمكنك السماح لـ `chai-http` بتشغيل وإيقاف الخادم تلقائيًا.

مع ذلك، تتطلب الاختبارات التي تُشغل عند إرسال رابط مشروعك أن يكون الخادم قيد التشغيل، لذا تحتاج إلى استخدام طريقة `keepOpen` لمنع `chai-http` من إيقاف الخادم.

# --instructions--

داخل `tests/2_functional-tests.js`، عدل اختبار `'Test GET /hello with no name'` (`// #1`) للتحقق من `status` و`text` في الاستجابة لجعل الاختبار ينجح. لا تغير المعلمات الممررة إلى الـ asserts.

يجب ألا يكون هناك استعلام في عنوان URL. بدون استعلام اسم في عنوان URL، تستجيب نقطة النهاية بـ `hello Guest`.

# --hints--

يجب أن تنجح جميع الاختبارات

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

يجب أن تختبر أن `res.status` == 200

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

يجب أن تختبر أن `res.text` == `'hello Guest'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.text');
assert.match(data.assertions[1].args[1], /('|")hello Guest\1/);
```

