---
id: 587d8250367417b2b2512c5d
title: تشغيل اختبارات وظيفية باستخدام متصفح بدون واجهة
challengeType: 2
forumTopicId: 301595
dashedName: run-functional-tests-using-a-headless-browser
---

# --description--

كتذكير، يُبنى هذا المشروع على مشروع بداية تم استنساخه من <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

توجد في الصفحة نموذج إدخال. يرسل بيانات إلى نقطة النهاية `PUT /travellers` كطلب AJAX.

عند اكتمال الطلب بنجاح، تُضيف شفرة العميل `<div>` يحتوي على المعلومات الموجودة في الاستجابة إلى DOM.

إليك مثالًا على كيفية استخدام Zombie.js للتفاعل مع النموذج:

```js
test('Submit the surname "Polo" in the HTML form', function (done) {
  browser.fill('surname', 'Polo').then(() => {
    browser.pressButton('submit', () => {
      browser.assert.success();
      browser.assert.text('span#name', 'Marco');
      browser.assert.text('span#surname', 'Polo');
      browser.assert.elements('span#dates', 1);
      done();
    });
  });
});
```

أولًا، تستخدم طريقة `fill` لكائن `browser` لملء حقل `surname` في النموذج بالقيمة `'Polo'`. تُرجع `fill` وعدًا، لذا يُتبع بـ `then`.

داخل رد النداء `then`، تُستخدم طريقة `pressButton` لكائن `browser` لاستدعاء مستمع حدث `submit` للنموذج. طريقة `pressButton` غير متزامنة.

ثم، بمجرد استلام استجابة من طلب AJAX، تُجرى بعض عمليات التحقق لتأكيد ما يلي:

1. حالة الاستجابة هي `200`
2. النص داخل العنصر `<span id='name'></span>` يطابق `'Marco'`
3. النص داخل العنصر `<span id='surname'></span>` يطابق `'Polo'`
4. يوجد عنصر واحد `1` من `<span id='dates'></span>`.

أخيرًا، يُستدعى رد النداء `done`، وهو ضروري بسبب الاختبار غير المتزامن.

# --instructions--

داخل الملف `tests/2_functional-tests.js`، في اختبار `'Submit the surname "Colombo" in the HTML form'` (`// #5`)، قم بأتمتة التالي:

1. املأ النموذج باللقب `Colombo`
2. اضغط زر الإرسال

وفي رد النداء الخاص بـ `pressButton`:

1. تحقق أن الحالة هي OK `200`
2. تحقق أن النص داخل العنصر `span#name` هو `'Cristoforo'`
3. تحقق أن النص داخل العنصر `span#surname` هو `'Colombo'`
4. تحقق أن العنصر/العناصر `span#dates` موجودة وعددها `1`

لا تنسَ إزالة استدعاء `assert.fail()`.

# --hints--

يجب أن تجتاز جميع الاختبارات.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

يجب أن تتحقق من نجاح طلب المتصفح بدون واجهة.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'browser.success');
```

يجب أن تتحقق من أن النص داخل العنصر `span#name` هو `'Cristoforo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'browser.text');
assert.match(data.assertions[1].args[0], /('|")span#name\1/);
assert.match(data.assertions[1].args[1], /('|")Cristoforo\1/);
```

يجب أن تتحقق من أن النص داخل العنصر `span#surname` هو `'Colombo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'browser.text');
assert.match(data.assertions[2].args[0], /('|")span#surname\1/);
assert.match(data.assertions[2].args[1], /('|")Colombo\1/);
```

يجب أن تتحقق من وجود العنصر `span#dates` وأن عدده 1.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'browser.elements');
assert.match(data.assertions[3].args[0], /('|")span#dates\1/);
assert.equal(data.assertions[3].args[1], 1);
```

