---
id: 587d824a367417b2b2512c46
title: تعلّم كيف تعمل تحققات JavaScript
challengeType: 2
forumTopicId: 301589
dashedName: learn-how-javascript-assertions-work
---

# --description--

سيتطلب منك العمل على هذه التحديات كتابة شفرتك باستخدام إحدى الطرق التالية:

- استنساخ <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">مستودع GitHub هذا</a> وإكمال هذه التحديات محليًا.
- استخدم منشئ مواقع من اختيارك لإكمال المشروع. تأكد من دمج جميع الملفات من مستودع GitHub الخاص بنا.

# --instructions--

داخل `tests/1_unit-tests.js` تحت الاختبار الموسوم بـ `#1` في مجموعة `Basic Assertions`، غيّر كل `assert` إلى إما `assert.isNull` أو `assert.isNotNull` لجعل الاختبار ينجح (يجب أن يُقيّم إلى `true`). لا تغير المعلمات الممررة إلى عمليات التحقق.

# --hints--

يجب أن تنجح جميع الاختبارات.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

اختر الطريقة الصحيحة للتحقق الأول - `isNull` مقابل `isNotNull`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isNull', 'Null is null');
```

اختر الطريقة الصحيحة للتحقق الثاني - `isNull` مقابل `isNotNull`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'isNotNull', '1 is not null');
```

