---
id: 587d824f367417b2b2512c5b
title: تشغيل اختبارات وظيفية على استجابة API باستخدام Chai-HTTP IV - طريقة PUT
challengeType: 2
forumTopicId: 301591
dashedName: run-functional-tests-on-an-api-response-using-chai-http-iv---put-method
---

# --description--

كتذكير، يُبنى هذا المشروع على مشروع بداية تم استنساخه من <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

هذا التمرين مشابه للتمرين السابق.

الآن بعد أن تعرف كيفية اختبار طلب `PUT`، حان دورك لتفعيله من الصفر.

# --instructions--

داخل الملف `tests/2_functional-tests.js`، عدّل اختبار `'Send {surname: "da Verrazzano"}'` (`// #4`) واستخدم الطريقتين `put` و`send` لاختبار نقطة النهاية `'/travellers'`.

أرسل كائن JSON التالي مع طلب PUT الخاص بك:

```json
{
  "surname": "da Verrazzano"
}
```

تحقق من التالي داخل استدعاء الراجعة `request.end`:

1. يجب أن تكون قيمة `status` هي `200`
2. يجب أن يكون نوع `type` هو `application/json`
3. يجب أن تكون قيمة `body.name` هي `Giovanni`
4. يجب أن تكون قيمة `body.surname` هي `da Verrazzano`

اتبع ترتيب التحقق أعلاه - نعتمد عليه. تأكد أيضًا من إزالة `assert.fail()` عند الانتهاء.

# --hints--

يجب أن تجتاز جميع الاختبارات

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

يجب أن تختبر أن `res.status` تساوي 200

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

يجب أن تختبر أن `res.type` تساوي `'application/json'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.type');
assert.match(data.assertions[1].args[1], /('|")application\/json\1/);
```

يجب أن تختبر أن `res.body.name` تساوي `'Giovanni'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'equal');
assert.equal(data.assertions[2].args[0], 'res.body.name');
assert.match(data.assertions[2].args[1], /('|")Giovanni\1/);
```

يجب أن تختبر أن `res.body.surname` تساوي `'da Verrazzano'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'equal');
assert.equal(data.assertions[3].args[0], 'res.body.surname');
assert.match(data.assertions[3].args[1], /('|")da Verrazzano\1/);
```

