---
id: 587d824e367417b2b2512c55
title: اختبار إذا كان للكائن خاصية
challengeType: 2
forumTopicId: 301604
dashedName: test-if-an-object-has-a-property
---

# --description--

كتذكير، يُبنى هذا المشروع على مشروع بداية تم استنساخه من <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

تتحقق `property` من أن الكائن الفعلي يحتوي على خاصية معينة.

# --instructions--

داخل `tests/1_unit-tests.js` تحت الاختبار الموسوم بـ `#16` في مجموعة `Objects`، غيّر كل `assert` إلى إما `assert.property` أو `assert.notProperty` لجعل الاختبار ينجح (يجب أن يُقيّم إلى `true`). لا تغير المعلمات الممررة إلى عمليات التحقق.

# --hints--

يجب أن تنجح جميع الاختبارات.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

اختر الطريقة الصحيحة للتحقق الأول - `property` أم `notProperty`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notProperty',
  'A car has not wings'
);
```

اختر الطريقة الصحيحة للتحقق الثاني - `property` أم `notProperty`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'property',
  'planes have engines'
);
```

اختر الطريقة الصحيحة للتحقق الثالث - `property` أم `notProperty`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'property', 'Cars have wheels');
```

