---
id: 587d824d367417b2b2512c53
title: اختبار ما إذا كانت سلسلة نصية تحتوي على سلسلة فرعية
challengeType: 2
forumTopicId: 301597
dashedName: test-if-a-string-contains-a-substring
---

# --description--

كتذكير، يُبنى هذا المشروع على مشروع بداية تم نسخه من <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

تعمل الدالتان `include()` و `notInclude()` مع السلاسل النصية أيضًا! تتحقق `include()` من أن السلسلة النصية الفعلية تحتوي على السلسلة الفرعية المتوقعة.

# --instructions--

داخل الملف `tests/1_unit-tests.js` تحت الاختبار الموسوم بـ `#14` في مجموعة `Strings`، غيّر كل `assert` إلى إما `assert.include` أو `assert.notInclude` لجعل الاختبار ينجح (يجب أن يُقيّم إلى `true`). لا تغير المعلمات الممررة إلى عمليات التحقق.

# --hints--

يجب أن تنجح كل الاختبارات.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

اختر الطريقة الصحيحة للتحقق الأول - `include` أم `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'include',
  "'Arrow' contains 'row'..."
);
```

اختر الطريقة الصحيحة للتحقق الثاني - `include` أم `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'notInclude',
  "... a 'dart' doesn't contain a 'queue'"
);
```

