---
id: 587d824c367417b2b2512c4c
title: تحقق من المساواة العميقة باستخدام .deepEqual و .notDeepEqual
challengeType: 2
forumTopicId: 301587
dashedName: assert-deep-equality-with--deepequal-and--notdeepequal
---

# --description--

كتذكير، يُبنى هذا المشروع على مشروع بداية تم نسخه من <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

تتحقق الدالة `deepEqual()` من أن كائنين متساويان بالمساواة العميقة.

# --instructions--

داخل الملف `tests/1_unit-tests.js` تحت الاختبار الموسوم بـ `#7` في مجموعة `Equality`، غيّر كل `assert` إلى إما `assert.deepEqual` أو `assert.notDeepEqual` لجعل الاختبار ينجح (يجب أن يُقيّم إلى `true`). لا تغير المعلمات الممررة إلى عمليات التحقق.

# --hints--

يجب أن تنجح جميع الاختبارات.

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
})();
```

اختر الطريقة الصحيحة للتحقق الأول - `deepEqual` أم `notDeepEqual`.

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'deepEqual',
    'The order of the keys does not matter'
  );
})();
```

اختر الطريقة الصحيحة للتحقق الثاني - `deepEqual` أم `notDeepEqual`.

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'notDeepEqual',
    'The position of elements within an array does matter'
  );
})();
```

