---
id: 587d824f367417b2b2512c5c
title: محاكاة الإجراءات باستخدام متصفح بدون واجهة
challengeType: 2
dashedName: simulate-actions-using-a-headless-browser
---

# --description--

كتذكير، يُبنى هذا المشروع على مشروع بداية تم استنساخه من <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

في التحديات التالية، ستحاكي التفاعل البشري مع صفحة باستخدام متصفح بدون واجهة.

المتصفحات بدون واجهة هي متصفحات ويب بدون واجهة مستخدم رسومية. يمكنها عرض وتفسير HTML وCSS وJavaScript بنفس طريقة المتصفح العادي، مما يجعلها مفيدة جدًا لاختبار صفحات الويب.

في التحديات التالية ستستخدم Zombie.js، وهو متصفح بدون واجهة خفيف الوزن لا يعتمد على تثبيت ملفات ثنائية إضافية. لكن هناك خيارات أخرى كثيرة لمتصفحات بدون واجهة أكثر قوة.

تسمح Mocha بتشغيل بعض الكود قبل بدء أي من الاختبارات الفعلية. هذا مفيد لإضافة مدخلات إلى قاعدة بيانات تُستخدم في بقية الاختبارات.

مع المتصفح بدون واجهة، قبل تشغيل الاختبارات، يجب عليك **زيارة** الصفحة التي ستختبرها.

يُنفذ الخطاف `suiteSetup` مرة واحدة فقط في بداية مجموعة الاختبارات.

هناك عدة أنواع أخرى من الـ hooks يمكنها تنفيذ كود قبل كل اختبار، بعد كل اختبار، أو في نهاية مجموعة الاختبارات. اطلع على وثائق Mocha لمزيد من المعلومات.

# --instructions--

داخل `tests/2_functional-tests.js`، مباشرة بعد إعلان `Browser`، أضف عنوان URL الخاص بمشروعك إلى خاصية `site` للمتغير:

```js
Browser.site = 'http://0.0.0.0:3000'; // Your URL here
```

ثم على المستوى الجذري لمجموعة `'Functional Tests with Zombie.js'`، أنشئ نسخة جديدة من كائن `Browser` باستخدام الكود التالي:

```js
const browser = new Browser();
```

واستخدم الـ `suiteSetup` hook لتوجيه `browser` إلى مسار `/` بالكود التالي. **ملاحظة**: يتم تمرير `done` كاستدعاء راجع إلى `browser.visit`، فلا يجب عليك استدعاؤه.

```js
suiteSetup(function(done) {
  return browser.visit('/', done);
});
```

# --hints--

يجب أن تجتاز جميع الاختبارات.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=4');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

