---
id: 587d824c367417b2b2512c4d
title: قارن خصائص عنصرين
challengeType: 2
forumTopicId: 301588
dashedName: compare-the-properties-of-two-elements
---

# --description--

كتذكير، يُبنى هذا المشروع على مشروع بداية تم نسخه من <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

داخل `tests/1_unit-tests.js` تحت الاختبار الموسوم بـ `#8` في مجموعة `Comparisons`، غيّر كل `assert` إلى إما `assert.isAbove` أو `assert.isAtMost` لجعل الاختبار ينجح (يجب أن يُقيّم إلى `true`). لا تغير المعلمات الممررة إلى عمليات التحقق.

# --hints--

يجب أن تنجح كل الاختبارات.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

اختر الطريقة الصحيحة للتحقق الأول - `isAbove` مقابل `isAtMost`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isAtMost',
  '5 is at most (<=) 5'
);
```

اختر الطريقة الصحيحة للتحقق الثاني - `isAbove` مقابل `isAtMost`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'isAbove', '1 is greater than 0');
```

اختر الطريقة الصحيحة للتحقق الثالث - `isAbove` مقابل `isAtMost`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[2].method,
  'isAbove',
  'Math.PI = 3.14159265 is greater than 3'
);
```

اختر الطريقة الصحيحة للتحقق الرابع - `isAbove` مقابل `isAtMost`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'isAtMost',
  '1 - Math.random() is > 0 and <= 1. It is atMost 1 !'
);
```

