---
id: 587d824d367417b2b2512c50
title: اختبار ما إذا كانت القيمة مصفوفة
challengeType: 2
forumTopicId: 301600
dashedName: test-if-a-value-is-an-array
---

# --description--

كتذكير، يُبنى هذا المشروع على مشروع بداية تم نسخه من <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

داخل `tests/1_unit-tests.js` تحت الاختبار الموسوم بـ `#11` في مجموعة `Arrays`، غيّر كل `assert` إلى إما `assert.isArray` أو `assert.isNotArray` لجعل الاختبار ينجح (يجب أن يُقيّم إلى `true`). لا تغير المعلمات الممررة إلى عمليات التحقق.

# --hints--

يجب أن تنجح جميع الاختبارات.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=10');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

اختر الطريقة الصحيحة للتحقق الأول - `isArray` مقابل `isNotArray`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=10');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isArray',
  'String.prototype.split() returns an Array'
);
```

اختر الطريقة الصحيحة للتحقق الثاني - `isArray` مقابل `isNotArray`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=10');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'isNotArray',
  'Array.prototype.indexOf() returns a number'
);
```

