---
id: 587d824c367417b2b2512c4e
title: اختبار ما إذا كانت قيمة واحدة أقل أو على الأقل مساوية لقيمة أخرى
challengeType: 2
forumTopicId: 301606
dashedName: test-if-one-value-is-below-or-at-least-as-large-as-another
---

# --description--

كتذكير، يُبنى هذا المشروع على مشروع بداية تم استنساخه من <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

داخل `tests/1_unit-tests.js` تحت الاختبار الموسوم بـ `#9` في مجموعة `Comparisons`، غيّر كل `assert` إلى إما `assert.isBelow` أو `assert.isAtLeast` لجعل الاختبار ينجح (يجب أن يُقيّم إلى `true`). لا تغير المعلمات الممررة إلى عمليات التحقق.

# --hints--

يجب أن تنجح كل الاختبارات.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

اختر الطريقة الصحيحة للتحقق الأول - `isBelow` مقابل `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isAtLeast',
  '5 is at least (>=) 5'
);
```

اختر الطريقة الصحيحة للتحقق الثاني - `isBelow` مقابل `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'isAtLeast',
  '2 * Math.random() is at least 0'
);
```

اختر الطريقة الصحيحة للتحقق الثالث - `isBelow` مقابل `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'isBelow', '1 is smaller than 2');
```

اختر الطريقة الصحيحة للتحقق الرابع - `isBelow` مقابل `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'isBelow',
  '2/3 (0.6666) is smaller than 1'
);
```

