---
id: 587d824c367417b2b2512c4f
title: اختبار ما إذا كانت القيمة تقع ضمن نطاق محدد
challengeType: 2
forumTopicId: 301598
dashedName: test-if-a-value-falls-within-a-specific-range
---

# --description--

كتذكير، يُبنى هذا المشروع على مشروع بداية تم نسخه من <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

```javascript
.approximately(actual, expected, delta, [message])
```

يتحقق من أن `actual` تساوي `expected` ضمن نطاق +/- `delta`.

# --instructions--

داخل `tests/1_unit-tests.js` تحت الاختبار الموسوم بـ `#10` في مجموعة `Comparisons`، غيّر كل `assert` إلى `assert.approximately` لجعل الاختبار ينجح (يجب أن يُقيّم إلى `true`).

اختر الحد الأدنى للنطاق (المعلمة الثالثة) لجعل الاختبار ينجح دائمًا. يجب أن يكون أقل من 1.

# --hints--

يجب أن تنجح كل الاختبارات.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

اختر النطاق الصحيح للتحقق الأول - `approximately(actual, expected, range)`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'approximately');
assert.equal(
  data.assertions[0].args[2],
  0.5,
  "weirdNumbers(0.5) is in the range (0.5, 1.5]. It's within 1 +/- 0.5"
);
```

اختر النطاق الصحيح للتحقق الثاني - `approximately(actual, expected, range)`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'approximately');
assert.equal(
  data.assertions[1].args[2],
  0.8,
  "weirdNumbers(0.2) is in the range (0.2, 1.2]. It's within 1 +/- 0.8"
);
```

