---
id: 5f8884f4c46685731aabfc41
title: تشغيل اختبارات وظيفية باستخدام متصفح بدون واجهة II
challengeType: 2
forumTopicId: 301594
dashedName: run-functional-tests-using-a-headless-browser-ii
---

# --description--

كتذكير، يُبنى هذا المشروع على مشروع بداية تم استنساخه من <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

داخل `tests/2_functional-tests.js`، في اختبار `'Submit the surname "Vespucci" in the HTML form'` (`// #6`)، أتمتة التالي:

1. املأ النموذج باللقب `Vespucci`
2. اضغط على زر الإرسال

وفي داخل استدعاء راجع `pressButton`:

1. تحقق أن الحالة OK `200`
2. تحقق أن النص داخل العنصر `span#name` هو `'Amerigo'`
3. تحقق أن النص داخل العنصر `span#surname` هو `'Vespucci'`
4. تحقق أن العنصر/العناصر `span#dates` موجودة وعددها `1`

لا تنسَ إزالة استدعاء `assert.fail()`.

# --hints--

يجب أن تجتاز جميع الاختبارات.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

يجب أن تتحقق من نجاح طلب المتصفح بدون واجهة.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'browser.success');
```

يجب أن تتحقق من أن النص داخل العنصر `span#name` هو `'Amerigo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'browser.text');
assert.match(data.assertions[1].args[0], /('|")span#name\1/);
assert.match(data.assertions[1].args[1], /('|")Amerigo\1/);
```

يجب أن تتحقق من أن النص داخل العنصر `span#surname` هو `'Vespucci'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'browser.text');
assert.match(data.assertions[2].args[0], /('|")span#surname\1/);
assert.match(data.assertions[2].args[1], /('|")Vespucci\1/);
```

يجب أن تتحقق من وجود العنصر `span#dates` وأن عدده 1.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'browser.elements');
assert.match(data.assertions[3].args[0], /('|")span#dates\1/);
assert.equal(data.assertions[3].args[1], 1);
```

