---
id: 616965351e74d4689eb6de30
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

You can have multiple `meta` elements on a web page. Each `meta` element adds information about the page that cannot be expressed by other HTML elements.

Füge ein weiteres `meta`-Element innerhalb des `head` ein. Gib ihm ein`name`-Attribut, welches auf `"viewport"` gesetzt ist und ein`content`-Attribut, das auf `"width=device-width, initial-scale=1.0"` gesetzt ist, damit deine Seite auf allen Geräten gleich aussieht.

# --hints--

Du solltest zwei `meta`-Elemente haben.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 2);
```

Dein neues`meta`-Element sollte ein leeres Element sein, es sollte kein schließendes Tag `</meta>` haben.

```js
assert.notMatch(code, /<\/meta>/i);
```


Dein neues `meta`-Element sollte ein `name`-Attribut haben, dass auf `"viewport"` gesetzt ist und ein `content`-Attribut, das auf `"width=device-width, initial-scale=1.0"` gesetzt ist.

```js
const meta = [...document.querySelectorAll('meta')];
const contentAttrRegex = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/i;
const target = meta?.find(m => m?.getAttribute('name')?.toLowerCase()?.replace(/\s*/g, '') === 'viewport' && contentAttrRegex.test(m?.getAttribute('content')));
assert.exists(target);
```

Deine beiden `meta`-Elemente sollten innerhalb des `head`-Elements liegen.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('meta').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
