---
id: 5d822fd413a79914d39e98d3
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Вкладіть чотири елементи `div` у контейнер `.bb1`. Дайте їм класи `bb1a`, `bb1b`, `bb1c` і `bb1d` у такому порядку. Ця будівля матиме чотири розділи.

# --hints--

Ви маєте створити чотири нові елементи `div`.

```js
assert.equal(document.querySelectorAll('div')?.length, 6);
```

Одному з нових елементів `div` слід надати клас `bb1a`.

```js
assert.exists(document.querySelector('div.bb1a'));
```

Одному з нових елементів `div` слід надати клас `bb1b`.

```js
assert.exists(document.querySelector('div.bb1b'));
```

Одному з нових елементів `div` слід надати клас `bb1c`.

```js
assert.exists(document.querySelector('div.bb1c'));
```

Одному з нових елементів `div` слід надати клас `bb1d`.

```js
assert.exists(document.querySelector('div.bb1d'));
```

Розмістіть нові елементи `div` у правильному порядку.

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert.isTrue(__t('div.bb1a','bb1b') && __t('div.bb1b','bb1c') && __t('div.bb1c','bb1d'));
```

Розмістіть нові елементи `div` всередині елемента `.bb1`.

```js
assert.equal(document.querySelectorAll('div.bb1 > div')?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
--fcc-editable-region--
      <div class="bb1"></div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
    
```
