---
id: 5d822fd413a79914d39e98cd
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

У CSS ви можете націлитися на все за допомогою зірочки. Додайте кордон до всього, використовуючи селектор `*`, і задавши йому властивість `border` зі значенням `1px solid black`. Це трюк, який допомагає візуалізувати, де знаходяться елементи і який у них розмір. Пізніше ви це видалите.

# --hints--

Ви повинні використати селектор `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Ви повинні використати властивість `border` для стилізації всіх елементів.

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

Усі елементи повинні мати кордон `1px solid black`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
