---
id: 5d822fd413a79914d39e98cc
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Ласкаво просимо до проєкту «Панорама CSS-змінних»! Start by adding a `link` element that links your `styles.css` file within the `head` element.

# --hints--

Ваш код повинен мати елемент `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

Ви не повинні змінювати наявні теґи `head`. Переконайтеся, що не видалили кінцевий теґ.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

You should have one void `link` element.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

Ваш елемент `link` повинен бути в межах елемента `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Ваш елемент `link` повинен мати атрибут `rel` зі значенням `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Ваш елемент `link` повинен мати атрибут `href` зі значенням `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
