---
id: 646c48df8674cf2b91020ecb
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

У вашому елементі `head` додайте тег `meta` з атрибутом `charset`, значення якого `utf-8`. Також додайте елемент `title` з текстом `fCC Cat Painting`.

# --hints--

Ви повинні додати один елемент `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Ваш елемент `meta` має мати атрибут `charset`, встановлений у `utf-8`.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Ви повинні додати один елемент `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

Ваш елемент `title` має містити текст `fCC Cat Painting`. Зверніть увагу, що важливі правопис і регістр.

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
