---
id: 646c47867800472a4ed5d2ea
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Почніть із базової структури HTML. Додайте `DOCTYPE` референс `html` та `html` елемент із `lang` атрибутом, встановленим у `en`. Також додайте `head` та `body` елементи всередині `html` елемента.

# --hints--

У вас має бути `DOCTYPE` оголошення `html`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Ваше `DOCTYPE` оголошення має бути на початку вашого HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

У вас має бути відкриваючий `html` тег із `lang` атрибутом `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

У вас має бути закриваючий `html` тег.

```js
assert(code.match(/<\/html>/i));
```

У вас має бути відкриваючий `head` тег.

```js
assert(code.match(/<head>/i));
```

У вас має бути закриваючий `head` тег.

```js
assert(code.match(/<\/head>/i));
```

У вас має бути відкриваючий `body` тег.

```js
assert(code.match(/<body>/i));
```

У вас має бути закриваючий `body` тег.

```js
assert(code.match(/<\/body>/i));
```

Ваш `body` елемент має йти після `head` елемента.

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

Ваші `head` та `body` елементи мають бути всередині `html` елемента.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
