---
id: 646c48df8674cf2b91020ecc
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Додайте `link` елемент всередину вашого `head` елемента. Для цього `link` елемента встановіть атрибут `rel` у значення `stylesheet` та атрибут `href` у значення `./styles.css`.

# --hints--

У вас має бути `link` елемент.

```js
assert.match(code, /<link/)
```

Ваш `link` елемент повинен мати атрибути `rel` та `href`, встановлені у значення `stylesheet` та `./styles.css`.

```js
assert.match(code, /<link\s+(rel=("|')stylesheet\2\s+href=("|')\.\/styles\.css\3|href=("|')\.\/styles\.css\4\s+rel=("|')stylesheet\5)\s*\/?>/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```

```css

```
