---
id: 646ceb843412c74edee27a79
title: Крок 22
challengeType: 0
dashedName: step-22
---

# --description--

Тепер ви навчитеся CSS-хитрощі для малювання трикутників.

Це буде використано для малювання вух у кота.

Використовуючи селектор класу, надайте елементу `.cat-right-ear` властивості `height` і `width`, встановлені в `100px`. Встановіть `background-color` у `white`. Встановіть ліві та праві кордони в `35px solid blue`. Встановіть верхні та нижні кордони в `35px solid red`.

# --hints--

Ви повинні мати селектор `.cat-right-ear`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Ваш селектор `.cat-right-ear` повинен мати властивість `height`, встановлену в `100px`. Не забудьте додати крапку з комою.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height, '100px')
```

Ваш селектор `.cat-right-ear` повинен мати властивість `width`, встановлену в `100px`. Не забудьте додати крапку з комою.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width, '100px')
```

Ваш селектор `.cat-right-ear` повинен мати властивість `background-color`, встановлену в `white`. Не забудьте додати крапку з комою.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor, 'white')
```

Ваш селектор `.cat-right-ear` повинен мати властивість `border-left`, встановлену в `35px solid blue`. Не забудьте додати крапку з комою.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid blue')
```

Ваш селектор `.cat-right-ear` повинен мати властивість `border-right`, встановлену в `35px solid blue`. Не забудьте додати крапку з комою.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid blue')
```

Ваш селектор `.cat-right-ear` повинен мати властивість `border-top`, встановлену в `35px solid red`. Не забудьте додати крапку з комою.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop, '35px solid red')
```

Ваш селектор `.cat-right-ear` повинен мати властивість `border-bottom`, встановлену в `35px solid red`. Не забудьте додати крапку з комою.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom, '35px solid red')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
