---
id: 6476fc5cf14b276e6d04e82a
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Використайте селектор класу, щоб задати вашому `.box` елементу ширину `200px`, висоту `600px` та колір фону `#000`. Також задайте йому `position` зі значенням `absolute`, `top` зі значенням `800px` і `left` зі значенням `650px`. 

# --hints--

У вас має бути селектор `.box`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box'))
```

Ваш селектор `.box` має мати властивість `width`, встановлену в `200px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.width === '200px')
```

Ваш селектор `.box` має мати властивість `height`, встановлену в `600px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.height === '600px')
```

Ваш селектор `.box` має мати властивість `background-color`, встановлену в `#000`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.backgroundColor === 'rgb(0, 0, 0)')
```

Ваш селектор `.box` має мати властивість `position`, встановлену в `absolute`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.position === 'absolute')
```

Ваш селектор `.box` має мати властивість `left`, встановлену в `650px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.left === '650px')
```

Ваш селектор `.box` має мати властивість `top`, встановлену в `800px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.top === '800px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: relative;
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
