---
id: 6476f7a4827bcc61682f2347
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Наступна властивість position — це `absolute`. Коли ви використовуєте значення `absolute` для властивості `position`, елемент виводиться з нормального потоку документа, а його позиція визначається властивостями `top`, `right`, `bottom` і `left`.

Встановіть властивість `position` вашого елемента `.cat-head` у значення `absolute`, потім задайте властивості `top` і `left` будь-яким додатним значенням у пікселях.

<!-- **Примітка**: Ви можете експериментувати з властивостями `top`, `left`, `bottom` і `right`, але тест пройде лише для `top` від `300px` і left від `400px`. -->

# --hints--

Ваш селектор `.cat-head` повинен мати властивість `position`, встановлену в `absolute`. Переконайтеся, що ви додали крапку з комою.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

Ваш селектор `.cat-head` повинен мати властивість `top`, встановлену в будь-яке число у `px`. Переконайтеся, що ви додали крапку з комою.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/)
```

Ваш селектор `.cat-head` повинен мати властивість `left`, встановлену в будь-яке число у `px`. Переконайтеся, що ви додали крапку з комою.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
