---
id: 6676a8a8d9861319ee901a90
title: Крок 24
challengeType: 0
dashedName: step-24
---

# --description--

Тепер ви можете почати працювати над вухами вашого кота.

Очистіть свій код повторення, видаливши селектор `.cat-right-ear` та всі його властивості.

Використовуючи селектор класу, задайте елементу `.cat-left-ear` лівий і правий кордон по `35px solid transparent` кожен. Також встановіть нижній кордон у `70px solid #5e5e5e`.

# --hints--


Ваш селектор `.cat-right-ear` не повинен мати властивість `height`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height)
```

Ваш селектор `.cat-right-ear` не повинен мати властивість `width`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width)
```


Ваш селектор `.cat-right-ear` не повинен мати властивість `border-left`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft)
```

Ваш селектор `.cat-right-ear` не повинен мати властивість `border-right`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight)
```

Ваш селектор `.cat-right-ear` не повинен мати властивість `border-top`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop)
```

Ваш селектор `.cat-right-ear` не повинен мати властивість `border-bottom`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom)
```

У вас не має бути селектора `.cat-right-ear`.

```js
assert(!new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

У вас має бути селектор `.cat-left-ear`.

```js 
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear'))
```

Ваш селектор `.cat-left-ear` має мати властивість `border-left`, встановлену в `35px solid transparent`

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderLeft === '35px solid transparent')
```

Ваш селектор `.cat-left-ear` має мати властивість `border-right`, встановлену в `35px solid transparent`

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderRight === '35px solid transparent')
```

Ваш селектор `.cat-left-ear` має мати властивість `border-bottom`, встановлену в `70px solid #5e5e5e`

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderBottom === '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 0;
  width: 0;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-top: 35px solid transparent;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
