---
id: 66a207974c806a19d6607073
title: "Crea un'app per l'inventario dei libri"
challengeType: 25
dashedName: build-a-book-inventory-app
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un elemento `h1` con il testo `Book Inventory`.
1. Dovresti avere un elemento `table` con colonne chiamate `Title`, `Author`, `Category`, `Status` e `Rate`.
1. La tua tabella dovrebbe avere almeno quattro righe, la prima per le intestazioni delle colonne e le altre riempite con informazioni sui libri.
1. Ogni riga della tabella all'interno del corpo della tabella dovrebbe avere la classe `read`, `to-read` o `in-progress`.
1. Gli elementi `td` della colonna `Status` dovrebbero contenere un elemento `span` con la `class` di `status` che circonda il testo `Read`, `To Read` o `In Progress`, a seconda della classe di quella riga.
1. Gli elementi `td` della colonna `Rate` dovrebbero contenere un elemento `span` con la `class` di `rate` che avvolge tre elementi `span` vuoti.
1. Gli elementi `.rate` posti dentro righe `.read` dovrebbero avere una classe aggiuntiva con il valore di `one`, `two` o `three`, a seconda della valutazione personale. Questo valore dovrebbe venire dopo `rate`.
1. Dovresti creare tre selettori di attributo per mirare agli elementi con la classe `read`, `to-read` e `in-progress` e impostare la loro proprietà `background-image` usando un `linear-gradient` a tua scelta.
1. Dovresti impostare la proprietà `display` di ogni elemento `span` su `inline-block`.
1. Dovresti usare un selettore di attributo per mirare agli elementi `span` con la classe `status` che sono discendenti di elementi `tr` con la classe `to-read` e impostare le loro proprietà `border` e `background-image`.
1. Dovresti usare un selettore di attributo per mirare agli elementi `span` con la classe `status` che sono discendenti di elementi `tr` con la classe `read` e impostare le loro proprietà `border` e `background-image`.
1. Dovresti usare un selettore di attributo per mirare agli elementi `span` con la classe `status` che sono discendenti di elementi `tr` con la classe `in-progress` e impostare le loro proprietà `border` e `background-image`.
1. Dovresti usare un selettore di attributo per mirare agli elementi `span` con la classe `status` e agli elementi `span` con un valore di classe che inizia con `rate` e impostare le loro proprietà `height`, `width` e `padding`.
1. Dovresti usare un selettore di attributo per mirare agli elementi `span` che sono figli diretti di elementi `span` con un valore `class` che inizia con `rate` e impostare le loro proprietà `border`, `border-radius`, `margin`, `height`, `width` e `background-color`.
1. Dovresti usare un selettore di attributo per mirare al primo discendente di elementi `span` che hanno la parola `one` come parte del loro valore `class` e impostare la sua proprietà `background-image` usando un `linear-gradient`.
1. Dovresti usare un selettore di attributo per mirare ai primi due discendenti di elementi `span` che hanno la parola `two` come parte del loro valore `class` e impostare la loro proprietà `background-image` usando un `linear-gradient`.
1. Dovresti usare un selettore di attributo per mirare ai tre elementi `span` che sono discendenti di elementi `span` che hanno la parola `three` come parte del loro valore `class` e impostare la loro proprietà `background-image` usando un `linear-gradient`.

# --hints--

Dovresti avere un elemento `h1` con il testo `Book Inventory`.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'Book Inventory');
```

Dovresti avere solo un elemento `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Dovresti avere un elemento `table`.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

Dovresti avere un elemento `thead` e un elemento `tbody` dentro `table`.

```js
assert.lengthOf(document.querySelectorAll("table > thead"), 1);
assert.lengthOf(document.querySelectorAll("table > tbody"), 1);
```

Dentro `thead` dovrebbe esserci un `tr` con 5 elementi `th`.

```js
assert.lengthOf(document.querySelectorAll("table > thead > tr"), 1);
assert.lengthOf(document.querySelectorAll("table > thead > tr > th"), 5);
```

La tua prima colonna dovrebbe avere il testo `Title` come intestazione.

```js
assert.equal(document.querySelectorAll('th')[0]?.textContent.trim(), 'Title');
```

La tua seconda colonna dovrebbe avere il testo `Author` come intestazione.

```js
assert.equal(document.querySelectorAll('th')[1]?.textContent.trim(), 'Author');
```

La tua terza colonna dovrebbe avere il testo `Category` come intestazione.

```js
assert.equal(document.querySelectorAll('th')[2]?.textContent.trim(), 'Category');
```

La tua quarta colonna dovrebbe avere il testo `Status` come intestazione.

```js
assert.equal(document.querySelectorAll('th')[3]?.textContent.trim(), 'Status');
```

La tua quinta colonna dovrebbe avere il testo `Rate` come intestazione.

```js
assert.equal(document.querySelectorAll('th')[4]?.textContent.trim(), 'Rate');
```

La tua tabella dovrebbe avere almeno quattro righe.

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
```

Ogni riga dovrebbe avere sempre 5 colonne.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isNotEmpty(rows);
for (const row of rows) {
  assert.lengthOf(row.children, 5);
}
```

Ogni riga della tabella tranne quella delle intestazioni dovrebbe avere la classe `read`, `to-read` o `in-progress`.

```js
const rows = document.querySelectorAll('tr');
assert.isNotEmpty(rows);
for (let i = 1; i < rows.length; i++) {
  const classList = [...rows[i].classList];
  const currentProgress = classList[0]; 
  assert.oneOf(currentProgress, ['read', 'to-read', "in-progress"]); 
}
```

Gli elementi `td` della colonna `Status` dovrebbero contenere un elemento `span`.

```js
const statusColumnData = document.querySelectorAll('td:nth-child(4)');
assert.isNotEmpty(statusColumnData);
for (let e of statusColumnData) {
  assert.equal(e?.children[0]?.tagName, 'SPAN');
}
```

Ogni elemento `span` della colonna `Status` dovrebbe avere la classe `status`.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
assert.isNotEmpty(statusSpans);
for (let e of statusSpans) {
    assert.isTrue(e.classList.contains('status'));
}
```

Ogni elemento `.status` dovrebbe avere il testo `Read`, `To Read` o `In Progress`, a seconda della classe della sua riga.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
const rows = Array.from(document.querySelectorAll('tr')).slice(1);
assert.isNotEmpty(statusSpans);
for (let i = 0; i < rows.length; i++) {
  switch (statusSpans[i]?.innerText.trim()) {
        case 'Read':
            assert.isTrue(rows[i].classList.contains('read'));
            break;
        case 'To Read':
            assert.isTrue(rows[i].classList.contains('to-read'));
            break;
        case 'In Progress':
            assert.isTrue(rows[i].classList.contains('in-progress'));
            break;
        default:
            assert.fail(); 
    }
}
```

Gli elementi `td` della colonna `Rate` dovrebbero contenere un elemento `span`.

```js
const rateColumnData = document.querySelectorAll('tr td:last-child');
assert.isNotEmpty(rateColumnData);
for (let e of rateColumnData) {
    assert.equal(e.children[0]?.tagName, 'SPAN')
}
```

Ogni elemento `span` che è figlio diretto di un elemento `td` della colonna `Rate` dovrebbe avere la classe `rate` come prima classe.

```js
const rateSpans = document.querySelectorAll('tr td:last-child > span:first-child');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.classList[0], 'rate');
}
```

Ogni elemento `.rate` dovrebbe contenere tre elementi `span` vuoti.

```js
const rateSpans = document.getElementsByClassName('rate');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.children.length, 3);
    for (let child of e.children) {
        assert.equal(child.tagName, 'SPAN');
        assert.equal(child.innerText.length, 0);
    }
}
```

Gli elementi `.rate` posti dentro righe `.read` dovrebbero avere una classe aggiuntiva dopo la classe `rate` con il valore `one`, `two` o `three`, a seconda della valutazione personale.

```js
const readBooksRates = document.querySelectorAll('.read .rate');
assert.isNotEmpty(readBooksRates);
for (let e of readBooksRates) {
    assert.oneOf(e.classList[1], ['one', 'two', 'three']);
}
```

Dovresti avere un selettore di attributo per mirare alle righe che hanno la classe `read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"]'));
```

Dovresti usare un selettore di attributo per mirare alle righe che hanno la classe `read` e impostare la loro proprietà `background-image` a un gradiente lineare a tua scelta.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="read"]')?.backgroundImage, 'linear-gradient(');
```

Dovresti avere un selettore di attributo per mirare alle righe che hanno la classe `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]'));
```

Dovresti usare un selettore di attributo per mirare alle righe che hanno la classe `to-read` e impostare la loro proprietà `background-image` a un gradiente lineare a tua scelta.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]')?.backgroundImage, 'linear-gradient(');
```

Dovresti avere un selettore di attributo per mirare alle righe che hanno la classe `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]'));
```

Dovresti usare un selettore di attributo per mirare alle righe che hanno la classe `in-progress` e impostare la loro proprietà `background-image` a un gradiente lineare a tua scelta.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]')?.backgroundImage, 'linear-gradient(');
```

Dovresti impostare la proprietà `display` di ogni elemento `span` su `inline-block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('span')?.getPropVal('display'), 'inline-block');
```

Dovresti avere un selettore di attributo per mirare agli elementi `span` con la classe `status` che sono discendenti di elementi `tr` con la classe `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]'));
```

Dovresti usare un selettore di attributo per mirare agli elementi `span` con la classe `status` che sono discendenti di elementi `tr` con la classe `to-read` e impostare la loro proprietà `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.border);
```

Dovresti usare un selettore di attributo per mirare agli elementi `span` con la classe `status` che sono discendenti di elementi `tr` con la classe `to-read` e impostare la loro proprietà `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.backgroundImage);
```

Dovresti avere un selettore di attributo per mirare agli elementi `span` con la classe `status` che sono discendenti di elementi `tr` con la classe `read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]'));
```

Dovresti usare un selettore di attributo per mirare agli elementi `span` con la classe `status` che sono discendenti di elementi `tr` con la classe `read` e impostare la loro proprietà `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.border);
```

Dovresti usare un selettore di attributo per mirare agli elementi `span` con la classe `status` che sono discendenti di elementi `tr` con la classe `read` e impostare la loro proprietà `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.backgroundImage);
```

Dovresti avere un selettore di attributo per mirare agli elementi `span` con la classe `status` che sono discendenti di elementi `tr` con la classe `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]'));
```

Dovresti usare un selettore di attributo per mirare agli elementi `span` con la classe `status` che sono discendenti di elementi `tr` con la classe `in-progress` e impostare la loro proprietà `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.border);
```

Dovresti usare un selettore di attributo per mirare agli elementi `span` con la classe `status` che sono discendenti di elementi `tr` con la classe `in-progress` e impostare la loro proprietà `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.backgroundImage);
```

Dovresti avere un selettore di attributo per mirare agli elementi `span` con la classe `status` e agli elementi `span` con un valore di classe che inizia con `rate`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.exists(selector1 || selector2);
```

Dovresti usare un selettore di attributo per mirare agli elementi `span` con la classe `status` e agli elementi `span` con un valore di classe che inizia con `rate` e impostare la loro proprietà `height`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.height || selector2?.height);
```

Dovresti usare un selettore di attributo per mirare agli elementi `span` con la classe `status` e agli elementi `span` con un valore di classe che inizia con `rate` e impostare la loro proprietà `width`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.width || selector2?.width);
```

Dovresti usare un selettore di attributo per mirare agli elementi `span` con la classe `status` e agli elementi `span` con un valore di classe che inizia con `rate` e impostare la loro proprietà `padding`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.padding || selector2?.padding);
```

Dovresti avere un selettore di attributo per mirare agli elementi `span` che sono figli diretti di elementi `span` con un valore `class` che inizia con `rate`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span'));
```

Dovresti usare un selettore di attributo per mirare agli elementi `span` che sono figli diretti di elementi `span` con un valore `class` che inizia con `rate` e impostare la loro proprietà `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border'));
```

Dovresti usare un selettore di attributo per mirare agli elementi `span` che sono figli diretti di elementi `span` con un valore `class` che inizia con `rate` e impostare la loro proprietà `border-radius`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border-radius'));
```

Dovresti usare un selettore di attributo per mirare agli elementi `span` che sono figli diretti di elementi `span` con un valore `class` che inizia con `rate` e impostare la loro proprietà `margin`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('margin'));
```

Dovresti usare un selettore di attributo per mirare agli elementi `span` che sono figli diretti di elementi `span` con un valore `class` che inizia con `rate` e impostare la loro proprietà `height`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('height'));
```

Dovresti usare un selettore di attributo per mirare agli elementi `span` che sono figli diretti di elementi `span` con un valore `class` che inizia con `rate` e impostare la loro proprietà `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('width'));
```

Dovresti usare un selettore di attributo per mirare agli elementi `span` che sono figli diretti di elementi `span` con un valore `class` che inizia con `rate` e impostare la loro proprietà `background-color`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('background-color'));
```

Dovresti avere un selettore di attributo per mirare al primo discendente di elementi `span` che hanno la parola `one` come parte del loro valore `class`.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Dovresti usare un selettore di attributo per mirare al primo discendente di elementi `span` che hanno la parola `one` come parte del loro valore `class` e impostare la sua proprietà `background-image` usando un `linear-gradient`.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image').includes('linear-gradient('));
```

Dovresti avere un selettore di attributo per mirare ai primi due discendenti di elementi `span` che hanno la parola `two` come parte del loro valore `class`.

```js

const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Dovresti usare un selettore di attributo per mirare ai primi due discendenti di elementi `span` che hanno la parola `two` come parte del loro valore `class` e impostare la loro proprietà `background-image` usando un `linear-gradient`.

```js
const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]

assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.backgroundImage.includes('linear-gradient(')) 
```

Dovresti avere un selettore di attributo per mirare agli elementi `span` che sono discendenti di elementi `span` che hanno la parola `three` come parte del loro valore `class`.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.exists(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Dovresti usare un selettore di attributo per mirare agli elementi `span` che sono discendenti di elementi `span` che hanno la parola `three` come parte del loro valore `class` e impostare la loro proprietà `background-image` usando un `linear-gradient`.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.include(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image'), 'linear-gradient(');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main>
    <h1>Book Inventory</h1>
    <div class="table-container">
      <table>
        <thead>
          <tr>
            <th>Title</th>
            <th>Author</th>
            <th>Category</th>
            <th>Status</th>
            <th>Rate</th>
          </tr>
        </thead>
        <tbody>
          <tr class="read">
            <td>The Three Musketeers</td>
            <td>Alexandre Dumas</td>
            <td>Historical Novel</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate three">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="to-read">
            <td>The Plague</td>
            <td>Albert Camus</td>
            <td>Philosophical Novel</td>
            <td><span class="status">To Read</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>The Metamorphosis</td>
            <td>Franz Kafka</td>
            <td>Novella</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate one">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Dead Souls</td>
            <td>Nikolai Gogol</td>
            <td>Picaresque</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="in-progress">
            <td>Lord of the Flies</td>
            <td>William Golding</td>
            <td>Allegorical Novel</td>
            <td><span class="status">In Progress</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Do Androids Dream of Electric Sheep?</td>
            <td>Philip K. Dick</td>
            <td>Science Fiction</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
        </tbody>
      </table>
    </div>
  </main>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
    font-family: Arial, sans-serif;
}

.table-container {
  overflow-x: auto;
}

table {
    border-collapse: collapse;
    width: 100vw;
}

thead {
    background-image: linear-gradient(hsl(197, 92%, 77%, 0.60), hsl(197, 92%, 50%, 0.60));
}

th {
    padding: 0.4em;
}

td {
    text-align: center;
    padding: 0.3em;
}

tr[class="read"] {
    background-image: linear-gradient(hsl(120, 100%, 85%, 0.6), hsl(120, 100%, 65%, 0.6));
}

tr[class="to-read"] {
    background-image: linear-gradient(hsl(120, 20%, 95%, 0.6), hsl(120, 20%, 75%, 0.6));
}

tr[class="in-progress"] {
    background-image: linear-gradient(hsl(40, 100%, 85%, 0.60), hsl(40, 100%, 65%, 0.6));
}

span {
    display: inline-block;
}

span[class="status"] {
    border-radius: 20%/60%;
}

span[class="status"],
span[class^="rate"] {
    height: 1.85em;
    width: 7em;
    padding: 0.3em;
}

tr[class="to-read"] span[class="status"] {
    border: 0.1em solid hsl(5, 100%, 45%);
    background-image: linear-gradient(hsl(5, 100%, 75%), hsl(5, 100%, 50%));
}

tr[class="read"] span[class="status"] {
    border: 0.1em solid hsl(120, 75%, 45%);
    background-image: linear-gradient(hsl(120, 75%, 75%), hsl(120, 75%, 50%));
}

tr[class="in-progress"] span[class="status"] {
    border: 0.1em solid hsl(40, 90%, 40%);
    background-image: linear-gradient(hsl(40, 90%, 75%), hsl(40, 90%, 50%));
}

span[class^="rate"] > span {
    border: 0.1em solid hsl(0, 0%, 50%);
    border-radius: 50%;
    background-color: hsl(0, 15%, 95%);
    height: 100%;
    width: 20%;
    margin: 0.1em;
}

span[class~="three"] span {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="two"] :nth-child(1),
span[class~="two"] :nth-child(2) {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="one"] :first-child {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}
```
