---
id: 68507cadbf36aa089a84ce2d
title: реалізувати алгоритм луна
challengeType: 27
dashedName: implement-the-luhn-algorithm
---

# --description--

Алгоритм Луна, також відомий як алгоритм «модуль 10» або «mod 10», — це проста формула контрольної суми, яка використовується для перевірки різних ідентифікаційних номерів, наприклад, номерів кредитних карток. Ось кроки для перевірки числа за допомогою алгоритму Луна:

- Починаючи справа і виключаючи крайню праву цифру (контрольну цифру), подвоюйте значення кожної другої цифри.
- Якщо результат подвоєння цифри більший за `9`, складіть цифри цього результату, щоб отримати одну цифру. Альтернативно, можна відняти `9` від результату.
- Підсумуйте всі цифри, включно з контрольною цифрою.
- Якщо сума всіх цифр кратна `10`, тоді число є дійсним; інакше — недійсним.

Припустимо, у нас є число `453914881`. Кроки для перевірки його за алгоритмом Луна будуть такими:

 ```md 
 Номер рахунку      4   5   3   9   1   4   8   8   1 
 Подвоїти кожну другу  4  10   3  18   1   8   8  16   1 
 Скласти цифри двозначних чисел   4   1   3   9   1   8   8   7   1 
 ``` 

- Потім складіть усі числа, `4 + 1 + 3 + 9 + 1 + 8 + 8 + 7 + 1 = 42`.
- Оскільки `42` не є кратним `10`, число недійсне.

У цій лабораторній роботі ви створите валідатор кредитної картки, використовуючи алгоритм Луна.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви маєте визначити функцію з ім’ям `verify_card_number`, яка приймає рядок цифр (що представляє номер картки) і перевіряє, чи є він дійсним згідно з алгоритмом Луна.

2. Усередині функції `verify_card_number`:

   - Ви маєте обробити будь-які дефіси або пробіли, які можуть бути присутні в номері картки, що передається.
   - Повернути `VALID!`, якщо номер картки дійсний; інакше повернути `INVALID!`.

Після завершення проєкту ви повинні побачити такі повідомлення залежно від введених даних:

| Номер картки         | Повідомлення  |
| ------------------- | -------- |
| 453914889           | VALID!   |
| 4111-1111-1111-1111 | VALID!   |
| 1234 5678 9012 3456 | INVALID! |


# --hints--

Ви маєте мати функцію з ім’ям `verify_card_number`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('verify_card_number')
    `)
})
```

`verify_card_number('453914889')` має повертати `VALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914889') == 'VALID!'
    `)
})
```

`verify_card_number('4111-1111-1111-1111')` має повертати `VALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('4111-1111-1111-1111') == 'VALID!'
    `)
})
```

`verify_card_number('453914881')` має повертати `INVALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914881') == 'INVALID!'
    `)
})
```

`verify_card_number('1234 5678 9012 3456')` має повертати `INVALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('1234 5678 9012 3456') == 'INVALID!'
    `)
})
```

`verify_card_number` має повертати `VALID!`, коли викликається з дійсним номером кредитної картки.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6467') == 'VALID!'
    assert verify_card_number('6011 1111 1111 1117') == 'VALID!'
    assert verify_card_number('3782-822463-10005') == 'VALID!'
    assert verify_card_number('5555 5555 5555 4444') == 'VALID!'
    assert verify_card_number('3056 9309 0259 04') == 'VALID!'
    assert verify_card_number('6011000990139424173') == 'VALID!'
    `)
})
```

`verify_card_number` має повертати `INVALID!`, коли викликається з недійсним номером кредитної картки.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6466') == 'INVALID!'
    assert verify_card_number('6011 1111 1111 1116') == 'INVALID!'
    assert verify_card_number('3782 822463 10004') == 'INVALID!'
    assert verify_card_number('5555-5555-5555-4440') == 'INVALID!'
    assert verify_card_number('3056 9309 0259 05') == 'INVALID!'
    assert verify_card_number('6011000990139424175') == 'INVALID!'
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def verify_card_number(card_number):
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    sum_of_odd_digits = 0
    card_number_reversed = translated_card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            number = (number // 10) + (number % 10)
        sum_of_even_digits += number

    total = sum_of_odd_digits + sum_of_even_digits
    if total % 10 == 0:
        return 'VALID!'
    else:
        return 'INVALID!'

# Example usage
verify_card_number('4111-1111-4555-1142')
```
