---
id: a6b0bb188d873cb2c8729495
title: реалізувати конвертер HTML-сутностей
challengeType: 26
dashedName: implement-an-html-entity-converter
---

# --description--

Ця лабораторна робота присвячена конвертації спеціальних символів у рядку у відповідні HTML-сутності.

**Мета**: Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути `convertHTML` функція, яка приймає рядок як аргумент.
1. `convertHTML` функція має повертати новий рядок, конвертуючи спеціальні символи у рядку-аргументі у відповідні HTML-сутності.

   - `&` має бути конвертовано у `&amp;`.
   - `<` має бути конвертовано у `&lt;`.
   - `>` має бути конвертовано у `&gt;`.
   - `"` має бути конвертовано у `&quot;`.
   - `'` має бути конвертовано у `&apos;`.

# --hints--

У вас має бути `convertHTML` функція.

```js
assert.isFunction(convertHTML);
```

`convertHTML("Dolce & Gabbana")` має повертати рядок `Dolce &amp; Gabbana`.

```js
assert.match(convertHTML('Dolce & Gabbana'), /Dolce &amp; Gabbana/);
```

`convertHTML("Hamburgers < Pizza < Tacos")` має повертати рядок `Hamburgers &lt; Pizza &lt; Tacos`.

```js
assert.match(
  convertHTML('Hamburgers < Pizza < Tacos'),
  /Hamburgers &lt; Pizza &lt; Tacos/
);
```

`convertHTML("Sixty > twelve")` має повертати рядок `Sixty &gt; twelve`.

```js
assert.match(convertHTML('Sixty > twelve'), /Sixty &gt; twelve/);
```

`convertHTML('Stuff in "quotation marks"')` має повертати рядок `Stuff in &quot;quotation marks&quot;`.

```js
assert.match(
  convertHTML('Stuff in "quotation marks"'),
  /Stuff in &quot;quotation marks&quot;/
);
```

`convertHTML("Schindler's List")` має повертати рядок `Schindler&apos;s List`.

```js
assert.match(convertHTML("Schindler's List"), /Schindler&apos;s List/);
```

`convertHTML("<>")` має повертати рядок `&lt;&gt;`.

```js
assert.match(convertHTML('<>'), /&lt;&gt;/);
```

`convertHTML("abc")` має повертати рядок `abc`.

```js
assert.strictEqual(convertHTML('abc'), 'abc');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const map = {
  '&': '&amp;',
  '<': '&lt;',
  '>': '&gt;',
  '"': '&quot;',
  "'": '&apos;'
};

function convertHTML(str) {
  return str.replace(/[&<>"']/g, function(char) {
    return map[char];
  });
}
```
