---
id: 65ef1d5e3d2927c5e0f4997b
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

У Python ключове слово `is` перевіряє ідентичність об’єктів. Його використовують, щоб визначити, чи вказують дві змінні на один і той самий об’єкт у пам’яті. На відміну від `is`, оператор рівності (`==`) визначає, чи однакові значення двох об’єктів, незалежно від того, чи є вони одним і тим самим об’єктом у пам’яті.

Поза межами циклу for створіть інструкцію `if`, щоб перевірити, чи `root is None`. Якщо так, виведіть повідомлення `'Failed to converge within {max_iterations} iterations.'`. Не забудьте відформатувати повідомлення за допомогою f-рядка.

# --hints--

Ви повинні створити інструкцію `if`, яка перевіряє, чи `root is None`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].find_conditions()[0].is_equivalent("root is None")`));

    }
})
```

Ви повинні передати аргумент `f'Failed to converge within {max_iterations} iterations.'` у функцію `print`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].is_equivalent("if root is None:\\n    print(f'Failed to converge within {max_iterations} iterations.')")`));
       
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if abs(square_mid - square_target) < tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

--fcc-editable-region--

--fcc-editable-region--

```
