---
id: 65ef1d104c2be9c4d8e22075
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

Якщо умови `if` і `elif` не виконуються, значення `mid` буде більшим за `square_target`. У цьому випадку створіть умову `else` і призначте значення `mid` змінній `high`.

# --hints--

У вас має бути умова `else`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_conditions()[2]== _Node()`));
    }
})

```

Ви маєте призначити значення `mid` змінній `high` в умові `else`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_bodies()[2].is_equivalent("high = mid")`));
    }
})
```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if abs(square_mid - square_target) < tolerance:
                root = mid
                break

--fcc-editable-region--
            elif square_mid < square_target:
                low = mid

--fcc-editable-region--
```
