---
id: 65ef1be2b958cabf712296cd
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Тепер ви будете багаторазово звужувати інтервал, знаходячи середину поточного інтервалу та порівнюючи квадрат середини з цільовим значенням.

Для цього всередині блоку `else` створіть цикл `for`, який виконається до `max_iterations` разів.

Для циклу використайте функцію `range`, яка генерує послідовність чисел, над якою можна ітерувати. Синтаксис такий: `range(start, stop, step)`, де `start` — початкове ціле число (включно), `stop` — останнє ціле число (не включно), а `step` — різниця між числом і попереднім у послідовності.

Також використайте `_` як змінну циклу. `_` слугує заповнювачем і корисний, коли потрібно використати змінну, але її значення не потрібне.

# --hints--

Вам слід створити цикл `for`, щоб ітерувати над `range(max_iterations)`. Використайте `_` як змінну циклу.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].is_equivalent("for _ in range(max_iterations):\\n    pass")`));
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--

--fcc-editable-region--
```
