---
id: 65ef19425d1b27b6c930bae6
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Якщо число, для якого ви хочете знайти квадратний корінь, є від’ємним, код має викликати помилку, оскільки квадратний корінь від від’ємного числа не визначений у дійсних числах.

Видаліть інструкцію `pass` і створіть інструкцію `if`, щоб перевірити, чи `square_target` менше за `0`.

# --hints--

Ви маєте видалити ключове слово `pass`.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").has_pass()`))
    }
})
```

Ви маєте створити інструкцію `if`, щоб перевірити, чи `square_target < 0`.

```js
({
    test: () => {
       assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_if("square_target < 0").is_empty()`))

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance = 1e-7, max_iterations = 100):
    pass
--fcc-editable-region--
```
