---
id: 65ef1a50049cf9bada13266f
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Створіть інструкцію `elif`, щоб перевірити, чи дорівнює `square_target` `0`. Якщо так, призначте значення `0` змінній `root`. Також виведіть повідомлення `'The square root of {square_target} is 0'`. Не забудьте відформатувати повідомлення за допомогою f-рядка.

# --hints--

Ви повинні мати інструкцію `elif`, щоб перевірити умову `square_target == 0`.


```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[1].is_equivalent("square_target == 0")`));
    }
})

```

Ви повинні призначити значення `0` змінній `root` і передати аргумент `f'The square root of {square_target} is 0'` у виклик `print`.


```js

({
    test: () => 
    {        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[1].is_equivalent("root = 0\\nprint(f'The square root of {square_target} is 0')")`));
    }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    
--fcc-editable-region--
```
