---
id: 65ef181dad3d96b14a21a03e
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви знайдете приблизний квадратний корінь заданого числа, використовуючи метод ділення навпіл.

Метод ділення навпіл — це техніка для знаходження коренів дійснозначної функції. Він працює шляхом звуження інтервалу, у якому лежить квадратний корінь, доки не досягне значення в межах заданої точності.

Почніть зі створення функції з іменем `square_root_bisection`. Параметри поки що залиште порожніми.

# --hints--

Ви повинні визначити функцію `square_root_bisection` без параметрів.

```js
({
    test: () => {
        assert(runPython(`
        import inspect        
        inspect.isfunction(square_root_bisection)
        `))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
