---
id: 65ef190c6b51e9b5a5f7ed29
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Надайте функції `square_root_bisection` такі параметри:

- `square_target`: Число, для якого потрібно знайти квадратний корінь.
- `tolerance` (необов’язковий): Допустима різниця між квадратом приблизного значення кореня та фактичним значенням цілі (за замовчуванням `1e-7`). Допуск `1e-7` означає, що розв’язок буде точним у межах `0.0000001` від істинного значення і є хорошим вибором за замовчуванням, який балансує точність і продуктивність.
- `max_iterations` (необов’язковий): Максимальна кількість ітерацій для виконання (за замовчуванням `100`). Якщо метод не збігається в межах цього ліміту, вважайте, що розв’язок не знайдено.
  
# --hints--

Ваша функція має мати такі параметри: `square_target`, `tolerance = 1e-7` і `max_iterations = 100`. Порядок має значення.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("square_root_bisection").has_args("square_target, tolerance=1e-7, max_iterations=100")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection():
    pass
--fcc-editable-region--

```
