---
id: 65ef198fde24dfb7ff675b42
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Інструкція `raise` дозволяє примусово викликати певну виняткову ситуацію. Вона складається з ключового слова `raise`, за яким іде тип винятку, і дає змогу надати кастомне повідомлення про помилку:

```py
raise ValueError("Invalid value")
```

Коли наведений код запускається, виникає `ValueError`, і користувачу показується повідомлення `"Invalid value"`.

Якщо `square_target` менше за `0`, неможливо обчислити квадратний корінь із дійсним значенням. Тому викличте `ValueError` з повідомленням `'Square root of negative number is not defined in real numbers'`. Не забудьте видалити ключове слово `pass`.

# --hints--

Ви повинні видалити ключове слово `pass`.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].has_pass()`))
    }
})
```

Ви повинні викликати `ValueError` з повідомленням `'Square root of negative number is not defined in real numbers'` всередині тіла `if`.

```js
({
    test: () => {
        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].is_equivalent("raise ValueError('Square root of negative number is not defined in real numbers')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        pass

--fcc-editable-region--
```
