---
id: 66ba762af611169359d9d369
title: 신문 기사 작성하기
challengeType: 25
dashedName: build-a-newspaper-article
demoType: onClick
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. HTML 문서의 루트 `font-size`를 `24px`로 설정해야 합니다.
1. 모든 다른 요소를 포함하는 `newspaper` 클래스를 가진 요소가 있어야 합니다.
1. `.newspaper` 요소는 `font-size`인 `16px`와 `Open Sans` 폰트, 그리고 `sans-serif` 대체 폰트를 가져야 합니다.
1. `.newspaper` 요소 안에는 최소 일곱 개 이상의 요소가 있어야 합니다: 신문 이름용 `name` 클래스를 가진 요소 하나, 기사 날짜용 `date` 클래스를 가진 요소 하나, 제목용 `headline` 클래스를 가진 요소 하나, 부제목용 `sub-headline` 클래스를 가진 요소 하나, 작성자용 `author` 클래스를 가진 요소 하나, 그리고 각각 `text` 클래스를 가진 두 개의 단락이 있어야 합니다. 이 모든 요소는 기사 정보로 채워져야 합니다.
1. `.name` 요소는 루트 요소 `font-size`의 두 배 크기인 `font-size`를 가져야 하며 `Times New Roman` 폰트와 `serif` 대체 폰트를 사용해야 합니다.
1. `.name`과 `.author` 요소는 CSS를 사용해 모든 문자를 대문자로 만들어야 합니다.
1. `.headline` 요소는 부모 요소 `font-size`의 두 배 크기인 `font-size`를 가져야 하며 굵게 표시해야 합니다.
1. `.sub-headline` 요소는 `font-weight`인 `100`와 부모 요소 `font-size`의 `1.5` 배 크기인 `font-size`를 가져야 하며 이탤릭체여야 합니다.
1. `.author`은 CSS를 사용해 굵게 만들어야 합니다.
1. `.text` 요소는 `text-indent`인 `20px`를 가져야 합니다.
1. `.text` 요소는 부모 요소 `line-height`의 두 배 크기인 `font-size`를 가져야 합니다.
1. `.text` 요소의 첫 글자는 굵게 표시하고 부모 요소 `font-size`의 두 배 크기여야 합니다. 이를 위해 `::first-letter` 선택자를 사용하세요.

# --hints--

`html` 선택자가 `font-size`를 `24px`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.fontSize, '24px');
```

`newspaper` 클래스를 가진 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('.newspaper'));
```

`.newspaper` 선택자가 요소들의 `font-size`를 `16px`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontSize, '16px');
```

`.newspaper` 선택자가 요소들의 `font-family`를 `'Open Sans', sans-serif`로 설정해야 합니다.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontFamily, ['"Open Sans", sans-serif', 'Open Sans, sans-serif']);
```

`name` 요소 안에 `.newspaper` 클래스를 가진 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('.newspaper > .name'));
```

`.name` 요소에는 신문 이름이 있어야 합니다.

```js
assert.isAtLeast(document.querySelector('.newspaper > .name')?.innerText.length, 1);
```

`date` 요소 안에 `.newspaper` 클래스를 가진 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('.newspaper > .date'));
```

`.date` 요소에는 날짜가 있어야 합니다.

```js
assert.isAtLeast(document.querySelector('.newspaper > .date')?.innerText.length, 1);
```

`headline` 요소 안에 `.newspaper` 클래스를 가진 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('.newspaper > .headline'));
```

`.headline` 요소에는 기사 제목이 있어야 합니다.

```js
assert.isAtLeast(document.querySelector('.newspaper > .headline')?.innerText.length, 1);
```

`sub-headline` 요소 안에 `.newspaper` 클래스를 가진 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('.newspaper > .sub-headline'));
```

`.sub-headline` 요소에는 기사 부제목이 있어야 합니다.

```js
assert.isAtLeast(document.querySelector('.newspaper > .sub-headline')?.innerText.length, 1);
```

`author` 요소 안에 `.newspaper` 클래스를 가진 요소가 있어야 합니다.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .author').length, 1);
```

`.author` 요소에는 기사 작성자가 있어야 합니다.

```js
assert.isAtLeast(document.querySelector('.newspaper > .author')?.innerText.length, 1);
```

`text` 요소 안에 각각 `.newspaper` 클래스를 가진 두 개 이상의 단락 요소가 있어야 합니다.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .text').length, 2);
```

`.text` 요소에는 기사 본문이 있어야 합니다.

```js
const textEls = document.querySelectorAll('.newspaper > .text');
assert.isAtLeast(textEls.length, 1);
textEls.forEach(el => assert.isAtLeast(el.innerText.length, 1))
```

`.name` 선택자가 요소들의 `font-size`를 `2rem`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.name')?.fontSize, '2rem');
```

`.name` 선택자가 요소들의 `font-family`를 `'Times New Roman', serif`로 설정해야 합니다.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.name')?.fontFamily, ['"Times New Roman", serif', 'Times New Roman, serif']);
```

`.name` 요소는 `text-transform` 속성을 `uppercase`로 설정해야 합니다.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.name')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

`.author` 요소는 `text-transform` 속성을 `uppercase`로 설정해야 합니다.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.author')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

`.headline` 선택자가 요소들의 `font-size`를 `2em`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.headline')?.fontSize, '2em');
```

`.headline` 선택자가 요소들의 `font-weight`를 `bold`로 설정해야 합니다.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.headline')?.fontWeight, ['bold', '700']);
```

`.sub-headline` 선택자가 요소들의 `font-weight`를 `100`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontWeight, '100');
```

`.sub-headline` 선택자가 요소들의 `font-size`를 `1.5em`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontSize, '1.5em');
```

`.sub-headline` 선택자가 요소들의 `font-style`를 `italic`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontStyle, 'italic');
```

`.author` 선택자가 요소들의 `font-weight`를 `bold`로 설정해야 합니다.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.author')?.fontWeight, ['bold', '700']);
```

`.text` 선택자가 요소들의 `text-indent`를 `20px`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.textIndent, '20px');
```

`.text` 선택자가 요소들의 `line-height`를 `2em`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.lineHeight, '2em');
```

`.text::first-letter` 선택자가 요소들의 `font-weight`를 `bold`로 설정해야 합니다.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontWeight, ['bold', '700']);
```

`.text::first-letter` 선택자가 요소들의 `font-size`를 `2em`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontSize, '2em');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Article</title>
  </head>
  <body>

  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Daily Chuckles</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main class="newspaper">
    <h1 class="name">Daily Chuckles</h1>
    <p class="date">July 5, 2024</p>
    <h2 class="headline">Breaking: Grandma Edna Saves Earth</h2>
    <h3 class="sub-headline">Alien Invasion Foiled by Tech-Savvy Grandma's Wi-Fi Password</h3>
    <p class="author">By Jane Doe</p>
    <p class="text">In an extraordinary twist of fate, Grandma Edna found herself at the forefront of a potential crisis
      when her clever Wi-Fi security measures unexpectedly thwarted an alien invasion. As alien spacecraft descended
      upon the town, panic spread until Grandma Edna calmly intervened, resetting her router with a cryptic passphrase
      known only to her.</p>
    <p class="text">The aliens, encountering unforeseen technological barriers, struggled to breach Grandma Edna's
      fortified network. Frustrated and bewildered, they eventually retreated to reassess their invasion strategy,
      leaving residents both relieved and amazed at Grandma Edna's resourcefulness.</p>
  </main>
</body>

</html>
```

```css
html {
  font-size: 24px;
  background-color: MintCream;
  padding-top: 100px;
}

.newspaper {
  font-family: 'Open Sans', sans-serif;
  max-width: 800px;
  margin: 20px auto;
  padding: 20px;
  background-color: white;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
  font-size: 16px;
}

.name {
  font-family: 'Times New Roman', serif;
  font-size: 2rem;
  text-transform: uppercase;
  text-align: center;
  margin-bottom: 10px;
}

.date {
  font-size: 1em;
  text-align: center;
  margin-bottom: 20px;
}

.headline {
  font-size: 2em;
  font-weight: bold;
  line-height: 1.2;
  margin-bottom: 10px;
}

.sub-headline {
  font-size: 1.5em;
  font-style: italic;
  font-weight: 100;
  line-height: 1.4;
  margin-bottom: 15px;
}

.author {
  font-size: 1em;
  text-transform: uppercase;
  font-weight: bold;
  margin-bottom: 20px;
}

.text {
  text-indent: 20px;
  font-size: 1em;
  line-height: 2em;
  margin-bottom: 15px;
}

.text::first-letter {
  font-size: 2em;
  font-weight: bold;
}
```
