---
id: 66a37f37ef5823a313de8c26
title: 月の軌道を作成する
challengeType: 25
dashedName: build-a-moon-orbit
demoType: onClick
---

# --description--

このラボでは、HTMLとCSSを使って月が地球の周りを回るシンプルなアニメーションを作成します。地球がシステムの中心にあり、月がその周りを軌道運動します。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. ビューポートの高さを`100%`に設定し、`body`要素にフレックスボックスレイアウトを設定して、すべての要素をページの中央に揃える必要があります。
   
2. `div`クラスを持つ`space`が必要です。

3. `.space`要素の中に、順番に`div`と`earth`クラスを持つ2つの`orbit`要素がある必要があります。

4. `.orbit`要素の中に、`div`クラスを持つ別の`moon`がある必要があります。
   
5. `div`の`class`を持つ`space`要素はページの中央に配置され、幅と高さが`200px`である必要があります。

6. `div`の`class`を持つ`space`要素は`relative`の位置指定を使用する必要があります。

7. `.earth`要素は`absolute`の位置指定を使用する必要があります。`.earth`要素の中心を親要素の縦（top）と横（left）の両軸の中間点に配置してください。その後、`.earth`要素を自身の幅と高さの半分だけ戻して、親要素である`.space`内で中央に配置します。

8. `.earth`要素は幅と高さが`100px`である必要があります。

9. `.orbit`要素は幅と高さが`200px`である必要があります。

10. `.orbit`要素は`absolute`の位置指定を使用して配置する必要があります。その右下の角が`.space`要素の中心にあり、`transform`プロパティで縦横両軸に`-50%`だけ移動させてください。

11. 月が地球の周りを回る軌道は円形である必要があります。

12. `.moon`要素は`absolute`の位置指定を使用し、幅と高さが`30px`である必要があります。`.moon`要素は`.orbit`要素の上部に位置し、水平方向に中央に配置される必要があります。

13. `.moon`要素の水平方向の位置をさらに調整し、その幅の半分だけ左に移動させてください。

14. `.earth`セレクターは背景色と`border-radius`が`50%`である必要があります。

15. `.moon`セレクターは背景色と`border-radius`が`50%`である必要があります。

16. `@keyframes orbit`という名前のアニメーションを定義し、`.orbit`要素を中心に360度回転させる必要があります。このアニメーションを`.orbit`要素に適用し、継続時間は`5`秒、タイミング関数は線形、繰り返しは無限にしてください。

**注意:** スタイルシートをHTMLにリンクし、CSSを適用してください。

# --hints--

`body`要素の内容は、ビューポートの高さの`100%`に設定し、`display`を`flex`に、`justify-content`を`align-items`に、`center`に設定してページの中央に配置してください。

```js

assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.display, 'flex');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.justifyContent, 'center');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.alignItems, 'center');
```

`div`クラスを持つメインの`space`が必要です。

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'space');
```

`.space`要素は幅と高さが`200px`である必要があります。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.height, '200px');
```

`.space`要素は`position`プロパティを`relative`に設定してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.position, 'relative')
```

`div`の`earth`を持つ`div`の中に、`class`クラスを持つ`space`が必要です。

```js
assert.exists(document.querySelector('div.space div.earth'));
```

`.earth`要素は`.space`要素の最初の子である必要があります。

```js
assert.strictEqual(document.querySelector('.space > div')?.getAttribute('class'), 'earth');
```

`.earth`要素は幅と高さが`100px`である必要があります。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.width, '100px');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.height, '100px');
```

`.earth`要素は`absolute`の位置指定を使用してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.position, 'absolute');
```

`.earth`要素は`top`と`left`の位置を`50%`に設定してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.top, '50%') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.left, '50%');
```

`.earth`要素は`50%`プロパティを`transform`に設定して、`translate(-50%, -50%)`だけ戻し上に移動させてください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.transform, 'translate(-50%, -50%)');
```

`div`要素の中に`orbit`クラスを持つ`.space`が必要です。

```js
assert.exists(document.querySelector('div.space div.orbit'));
```

`.orbit`要素は`.earth`要素の後に配置してください。

```js
assert.isTrue(document.querySelector('.earth')?.nextElementSibling?.classList?.contains('orbit'));
```

`.orbit`要素は幅と高さが`200px`である必要があります。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.height, '200px');
```

`.orbit`要素は`absolute`の位置指定を使用してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.position, 'absolute');
```

`.orbit`要素は`.space`要素の中心に配置し、`transform`プロパティで縦横両軸に`-50%`だけ移動させてください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.transform, 'translate(-50%, -50%)');
```

`div`要素の中に`moon`クラスを持つ`.orbit`が必要です。

```js
assert.exists(document.querySelector('div.space div.orbit div.moon'));
```

`.moon`要素は幅と高さが`30px`である必要があります。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.width, '30px') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.height, '30px');
```

`.moon`要素は`absolute`の位置指定を使用して配置してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.position, 'absolute');
```

`.moon`要素の上端は、その包含ブロックの上端と揃える必要があります。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.top, '0px');
```

`.moon`要素の左端を包含ブロックの中心（左から`50%`）に移動させてください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.left, '50%');
```

`.moon`要素はx軸方向に`-50%`だけ移動させて水平方向に調整してください。

```js
const transformValue = new __helpers.CSSHelp(document).getStyle('.moon')?.transform;
const validTransforms = ['translateX(-50%)', 'translate(-50%)'];
assert.oneOf(transformValue, validTransforms);
```

`.earth`セレクターは背景色を持つ必要があります。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.earth')?.backgroundColor);
```

`.earth`セレクターは`border-radius`が`50%`である必要があります。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.borderRadius, '50%');
```

`.moon`セレクターは背景色を持つ必要があります。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.moon')?.backgroundColor);
```

`.moon`セレクターは`border-radius`が`50%`である必要があります。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.borderRadius, '50%');
```

`@keyframes`ルールを作成してください。

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('keyframes'), 1);
```

新しい`@keyframes`ルールは`orbit`という名前にしてください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name, 'orbit');
```

`@keyframes orbit`ルールは`0%`セレクターを持つ必要があります。

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules; 
assert.strictEqual(rules?.[0]?.keyText, '0%');
```

`@keyframes orbit`ルールは`100%`セレクターを持つ必要があります。

```js 
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules
assert.isTrue(rules?.[0]?.keyText === '100%' || rules?.[1]?.keyText === '100%');
```

`0%`セレクターは`transform`プロパティを`rotate(0deg) translate(-50%, -50%)`に設定してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[0]?.style?.transform, 'rotate(0deg) translate(-50%, -50%)');
```

`100%`セレクターは`0%`セレクターの後に配置してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.keyText, '100%')
```

`100%`セレクターは`transform`プロパティを`rotate(360deg) translate(-50%, -50%)`に設定してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.style?.transform, 'rotate(360deg) translate(-50%, -50%)')
```

`.orbit`セレクターは`animation`プロパティを`orbit 5s linear infinite`に設定してください。

```js

const orbitStyles = new __helpers.CSSHelp(document).getStyle('.orbit');
assert.strictEqual(orbitStyles?.animationName, "orbit");
assert.strictEqual(orbitStyles?.animationDuration, "5s");
assert.strictEqual(orbitStyles?.animationTimingFunction, "linear");
assert.strictEqual(orbitStyles?.animationIterationCount, "infinite");

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Moon Orbit</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Moon Orbiting Earth</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="space">
        <div class="earth"></div>
        <div class="orbit">
            <div class="moon"></div>
        </div>
    </div>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #000;
}

.space {
    position: relative;
    width: 200px;
    height: 200px;
}

.earth {
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    left: 50%;
    transform: translate(-50%, -50%);
    box-shadow: 0 0 20px rgba(0, 0, 255, 0.5);
}

.orbit {
    width: 200px;
    height: 200px;
    position: absolute;
    transform: translate(-50%, -50%);
    animation: orbit 5s linear infinite;
}

.moon {
    width: 30px;
    height: 30px;
    background-color: gray;
    border-radius: 50%;
    position: absolute;
    top: 0;
    left: 50%;
    transform: translateX(-50%);
    box-shadow: 0 0 10px rgb(255, 255, 255);
}


@keyframes orbit {
    0% {
        transform: rotate(0deg) translate(-50%, -50%);
    }

    100% {
        transform: rotate(360deg) translate(-50%, -50%);
    }
}
```
