---
id: 658225d560369446ece5518b
title: Paso 13
challengeType: 20
dashedName: step-13
---

# --description--

Agrega otro par clave-valor al diccionario al que estás agregando a la lista `expense`. Usa la cadena `'category'` como clave, y el parámetro `category` como valor.

# --hints--

Debes agregar un segundo par clave-valor al diccionario que estás pasando como argumento al método `.append()`.

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict) and len(node[0].tree.keys) == 2
`)) })
```

Debes pasar `{'amount': amount, 'category': category}` como argumento al método `.append()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount, 'category': category})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount})
--fcc-editable-region--

expenses = []
```
