---
id: 65822773f4c60a479ec89f66
title: Paso 15
challengeType: 20
dashedName: step-15
---

# --description--

Dentro de la función `print_expenses`, crea un bucle `for` que itere sobre cada elemento en la lista `expenses`. Usa `expense` como la variable del bucle y mueve `pass` dentro del cuerpo del bucle.

# --hints--

Debes crear un ciclo `for` que itere sobre `expenses` en tu función `print_expenses`. Recuerda usar `pass` dentro del cuerpo del ciclo.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_iter().is_equivalent("expenses")`)) })
```

Debes usar `expense` como la variable del ciclo.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_vars().is_equivalent("expense")`)) })
```

Debes tener `pass` solo dentro del cuerpo de tu ciclo.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    pass
--fcc-editable-region--

expenses = []
```
