---
id: 65823cfc74aa564ffc460489
title: Paso 26
challengeType: 20
dashedName: step-26
---

# --description--

Ahora, llama a `map()` pasando tu funci贸n `lambda` como el primer argumento y la lista `expenses` como el segundo argumento.

# --hints--

Debes llamar a la funci贸n `map()` dentro de la funci贸n `total_expenses`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("total_expenses").find_calls("map")) == 1`)) })
```

Debes pasar tu funci贸n `lambda` como el primer argumento a la llamada `map()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[0].is_equivalent("lambda expense: expense['amount']")`)) })
```

Debes pasar la lista `expenses` como el segundo argumento a la llamada `map()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    lambda expense: expense['amount']
--fcc-editable-region--

expenses = []
```
