---
id: 6582575b8089f85b8b92d7c8
title: Paso 42
challengeType: 20
dashedName: step-42
---

# --description--

Una vez que tengas los detalles del gasto, necesitas llamar a la función `add_expense` para agregar el nuevo gasto a la lista `expenses`.

Después de obtener el `amount` y `category` usando `input()`, llama a la función `add_expense`, pasando tres argumentos: `expenses`, `amount` y `category`. 

- `expenses` es la lista vacía creada en la función principal anteriormente en este proyecto.
- `amount` es la cantidad del gasto.
- `category` es la categoría del gasto.

# --hints--

Debes llamar a la función `add_expense()` dentro de tu sentencia `if`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_calls("add_expense")) == 1`)) })
```

Debes pasar tres argumentos a tu llamada `add_expense()`: `expenses`, `amount` y `category`. El orden importa.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("add_expense(expenses, amount, category)")`)) })
```

Tu llamada a `add_expense()` debe venir después de la asignación de `amount` y `category`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].is_equivalent("amount = float(input('Enter amount: '))\\ncategory = input('Enter category: ')\\nadd_expense(expenses, amount, category)")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            
--fcc-editable-region--
```
