---
id: 65822bd82d708c4895080c35
title: Paso 17
challengeType: 20
dashedName: step-17
---

# --description--

En Python, algo importante que debes saber es que el mismo tipo de comilla usada para definir una cadena no puede usarse dentro de ella. Por ejemplo, la cadena `'I'm a string!'` no es válida. Para usar la comilla simple dentro de esa cadena, debes:

- Escapa la comilla anteponiendo una barra invertida: `'I\'m a string!'`
- O usa comillas dobles para definir la cadena: `"I'm a string!"` (preferido).

Puedes acceder a los valores en un diccionario a través de sus claves. Necesitas usar la notación de corchetes e incluir la clave entre los corchetes cuadrados:

```py
my_dict = {'amount': 50.0, 'category': 'Food'}
my_dict['amount'] # 50.0
```

Actualmente estás interpolando el diccionario `expense` en tu f-string. Modifica la expresión del f-string para acceder al valor de la clave `'amount'` y la clave `'category'` en el diccionario `expense`.

# --hints--

Debes pasar `f'Amount: {expense["amount"]}, Category: {expense["category"]}'` a tu llamada `print()`. Recuerda usar comillas dobles dentro de tu f-string con comillas simples y viceversa.

```js
({ test: () => assert.match(code, /^\s+print\s*\(\s*f("|')Amount: \{\s*expense\s*\[\s*(?=[^\1])("|')amount\2\s*\]\s*\}, Category: \{\s*expense\s*\[\s*(?=[^\1])("|')category\3\s*\]\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense}, Category: {expense}')
--fcc-editable-region--

expenses = []
```
