---
id: 6582507654b3ed5712341382
title: Paso 38
challengeType: 20
dashedName: step-38
---

# --description--

Vas a usar sentencias condicionales para verificar la elección del usuario. Si la elección es `'1'`, significa que el usuario quiere agregar un gasto.

Todavía dentro del ciclo `while`, bajo la variable `choice`, escribe una estructura `if` para verificar si `choice` es igual a la cadena `'1'`. Si es verdadero, será el punto de partida para agregar un nuevo gasto.

Dentro del cuerpo de la sentencia `if`, declara una variable `amount` y asígnale una llamada vacía a `input()`.

# --hints--

Debes crear una estructura `if` que verifique si `choice` es igual a `'1'` en tu ciclo `while`.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
cond.is_equivalent("choice == '1'") or cond.is_equivalent("'1' == choice")
`)) })
```

Debes declarar una variable llamada `amount` en el cuerpo de tu sentencia `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("amount")
`)) })
```

Debes asignar una llamada vacía de `input()` a tu variable `amount`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("amount").is_equivalent("amount = input()")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')

--fcc-editable-region--
```
