---
id: 6582324b5d37fe4b622bc3c0
title: Paso 19
challengeType: 20
dashedName: step-19
---

# --description--

Las funciones lambda son funciones breves y anónimas en Python, ideales para tareas simples y de una sola vez. Se definen con la palabra clave `lambda`, y usan la siguiente sintaxis:

```py
lambda x: expr
```

En el ejemplo anterior, `x` representa un parámetro que se usa en la expresión `expr`, y funciona igual que cualquier parámetro en una función tradicional. `expr` es la expresión que se evalúa y se retorna cuando se llama a la función lambda.

Crea una variable llamada `test` y asígnale una función lambda que tome un parámetro `x` y devuelva `x * 2`.

# --hints--

Debes declarar una variable llamada `test`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("test")`))})
```

Debes asignar `lambda x: x * 2` a tu variable `test`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("test").is_equivalent("test = lambda x: x * 2")`))})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass

--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
