---
id: 65824111a09164518320088d
title: Paso 30
challengeType: 20
dashedName: step-30
---

# --description--

La función `filter()` te permite seleccionar elementos de un iterable, como una lista, basándote en la salida de una función: 

```py
filter(my_function, my_list)
```

`filter()` toma una función como su primer argumento y un iterable como su segundo argumento. Devuelve un iterador, que es un objeto especial que te permite iterar sobre los elementos de una colección, como una lista.

El resultado del ejemplo anterior es un iterador que contiene los elementos de `my_list` para los cuales `my_function` devuelve `True`. 

Dentro de la función `filter_expenses_by_category`, llama a `filter()` pasando la función `lambda` que escribiste en el paso anterior como el primer argumento y la lista `expenses` como el segundo argumento.

# --hints--

Debes llamar a `filter()` dentro de la función `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")) == 1`)) })
```

Debes pasar `lambda expense: expense['category'] == category` como el primer argumento a la llamada `filter()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[0].is_equivalent("lambda expense: expense['category'] == category")`)) })
```

Debes pasar `expenses` como el segundo argumento a la llamada `filter()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    lambda expense: expense['category'] == category
--fcc-editable-region--

expenses = []
```
