---
id: 65823040bdb32949d4c1a96f
title: Paso 18
challengeType: 20
dashedName: step-18
---

# --description--

Necesitarás una función para calcular el monto total de los gastos.

Define una función llamada `total_expenses` que tome un parámetro `expenses`. Llena el cuerpo de la función con una sentencia `pass` por ahora.

# --hints--

Debes definir una función llamada `total_expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(total_expenses)
  `))
})
```

Tu función `total_expenses` debe recibir un solo parámetro `expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(total_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
