---
id: 65823dde36392f505a39f7c7
title: Paso 27
challengeType: 20
dashedName: step-27
---

# --description--

Finalmente, pase su `map()` llamada a la funci칩n `sum()` para obtener los gastos totales y devolver el resultado.

# --hints--

Debes llamar a la funci칩n `sum()` pas치ndole tu actual mapeo `` como argumento y devolver el resultado de tu funci칩n `total_expenses`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").has_return("sum(map(lambda expense: expense['amount'], expenses))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    map(lambda expense: expense['amount'], expenses)
--fcc-editable-region--

expenses = []
```
