---
id: 65822308ded1ba4632f66aa3
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

Un diccionario es otro tipo de dato incorporado en Python. Un diccionario es una colección de datos en forma de pares *clave*-*valor*. Los diccionarios se definen con llaves (`{}`) y contienen pares clave-valor separados por comas. Cada clave va seguida de dos puntos (`:`) y el valor:

```py
{'amount': 50.0, 'category': 'Food'}
```

En el ejemplo anterior, `'amount'` y `'category'` son claves, y `50.0` y `'Food'` son sus valores correspondientes.

Crea un diccionario con una clave `'amount'` y un valor del parámetro `amount` y pasa tu nuevo diccionario a la llamada `.append()`.

# --hints--

Debes pasar un diccionario como argumento al método `.append()`.

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict)
`)) })
```

Debes pasar `{'amount': amount}` como argumento al método `.append()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append()
--fcc-editable-region--

expenses = []
```
