---
id: 65823e87c9741750a22085a7
title: Paso 28
challengeType: 20
dashedName: step-28
---

# --description--

A continuación, define una función llamada `filter_expenses_by_category` que recibe dos parámetros: `expenses` y `category`. Usa `pass` para completar el cuerpo de la función.

# --hints--

Debes definir una función llamada `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(filter_expenses_by_category)
  `))
})
```

Tu función `filter_expenses_by_category` debe tomar `expenses` y `category` como parámetros.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(filter_expenses_by_category))
    sig == '(expenses, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
