---
id: 658212ba98182f3e855e85f9
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

En este proyecto, vas a construir un rastreador de gastos simple, pero funcional en Python.

Comienza definiendo una función llamada `add_expense` que recibe tres parámetros: `expenses`, `amount` y `category`. Usa la palabra clave `pass` para completar el cuerpo de la función.

# --hints--

Debes usar la palabra clave `def` para declarar una función llamada `add_expense`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(add_expense)
  `))
})
```

Tu función `add_expense` debe tomar tres parámetros: `expenses`, `amount` y `category`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(add_expense))
    sig == '(expenses, amount, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
