---
id: 658238f7604f154ea9a23e1e
title: Pasos 23
challengeType: 20
dashedName: step-23
---

# --description--

La función `sum()` devuelve la suma de los elementos en el iterable que se pasa como argumento. Vas a usar `sum()` junto con las funciones `map()` y `lambda` para obtener el monto total de los gastos.

Por ahora, haz una pequeña prueba y modifica tu llamada actual a `print()` reemplazando la llamada a `list()` con una llamada a la función `sum()` pasando como argumento la llamada actual a `map()`.

# --hints--

Debes imprimir `sum(map(test, [2, 3, 5, 8]))`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*sum\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(list(map(test, [2,3,5,8])))
--fcc-editable-region--

expenses = []
```
