---
id: 65822934c8a0454803e3ef12
title: Paso 16
challengeType: 20
dashedName: step-16
---

# --description--

A continuación, vas a mostrar los detalles de cada gasto.

Dentro del ciclo `for`, reemplaza `pass` con una llamada a `print()` y pásale la siguiente f-string: `f'Amount: {expense}, Category: {expense}'`.

# --hints--

Debes imprimir `f'Amount: {expense}, Category: {expense}'` en tu ciclo `for`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].has_call("print(f'Amount: {expense}, Category: {expense}')")`)) })
```

No debes tener `pass` dentro del cuerpo de tu ciclo.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_bodies()[0].has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        pass
--fcc-editable-region--

expenses = []
```
