---
id: 65824561f3478e5371a33ae5
title: Paso 32
challengeType: 20
dashedName: step-32
---

# --description--

El siguiente paso es definir la función principal, que será el punto de entrada del programa interactivo de seguimiento de gastos.

Define una función llamada `main` sin parámetros. Llena el cuerpo de la función con la lista `expenses` que creaste al inicio de este proyecto. Usarás esta lista para almacenar los registros de gastos.

# --hints--

Debes definir una función llamada `main()` sin parámetros.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(main)
    sig = str(inspect.signature(main))
    sig == '()'
  `))
})
```

Debes mover la lista `expenses` dentro del cuerpo de la función `main()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_body().is_equivalent("expenses = []")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--

expenses = []
--fcc-editable-region--
```
