---
id: 65821c1c40ccac44adf38e6b
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

Crea una lista vacía llamada `expenses`. La usarás para almacenar cada uno de tus gastos.

# --hints--

Debes declarar una nueva variable llamada `expenses`.

```js
({ test: () => assert(__userGlobals.has("expenses")) })
```

Debes asignar una lista vacía a la variable `expenses`.

```js
({ test: () => assert.equal(__userGlobals.get("expenses").length, 0) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    pass
    
--fcc-editable-region--

--fcc-editable-region--
```
