---
id: 65823ff0d4b991510fade1a8
title: Paso 29
challengeType: 20
dashedName: step-29
---

# --description--

Dentro de la función `filter_expenses_by_category`, reemplaza `pass` con una función `lambda`. Usa `expense` como parámetro y evalúa la comparación entre el valor de la clave `'category'` del diccionario `expense` y `category` usando el operador de igualdad.

# --hints--

Debes crear una función `lambda` que use un parámetro llamado `expense` y evalúe la expresión `expense['category'] == category` dentro de la función `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_stmt("lambda expense: expense['category'] == category")`)) })
```

No debes tener `pass` dentro de la función `filter_expenses_by_category`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    pass
--fcc-editable-region--

expenses = []
```
