---
id: 658246c28575d653d1f89b59
title: Paso 33
challengeType: 20
dashedName: step-33
---

# --description--

Un bucle `while` es otro tipo de bucle que ejecuta una porción de código mientras una condición especificada sea `True`. El bucle termina cuando la condición se vuelve `False`:

```py
while condition:
    <code>
```

Debajo de la lista `expenses`, crea un ciclo `while`. Usa `True` para la condición y muestra la cadena `'\nExpense Tracker'` dentro del cuerpo del ciclo para mostrar el título del programa.

# --hints--

Debes crear un ciclo `while` usando `True` como condición dentro de tu función `main()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_conditions()[0].is_equivalent("True")`)) })
```

Debes imprimir `'\nExpense Tracker'` en tu ciclo `while`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")[0].is_equivalent("print('\\\\nExpense Tracker')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--
def main():
    expenses = []
    
--fcc-editable-region--
```
