---
id: 6582352cbdbcde4d0c4f7b0b
title: Paso 20
challengeType: 20
dashedName: step-20
---

# --description--

Para llamar a una función lambda puedes usar la sintaxis habitual de función con un par de paréntesis después del nombre de la variable.

Llama a tu función lambda `test` y pasa `3` como argumento. Luego, imprime el resultado.

# --hints--

Debes tener `print(test(3))` en tu código.

```js
({ test: () => assert.match(code, /^print\s*\(\s*test\s*\(\s*3\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2

--fcc-editable-region--

expenses = []
```
