---
id: 65821fcc010c3245718f2a06
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

El parámetro `expenses` de tu función `add_expense` será una lista de gastos. Quieres poder agregar elementos al final de tu lista. Para eso usarás el método de lista `.append()`.

Agrega una llamada al método `.append()` en la lista `expenses`. No pases ningún argumento a `.append()` por ahora.

# --hints--

Debes tener `expenses.append()` en tu función.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append()")`)) })
```

No debes tener `pass` en tu función.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("add_expense").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    pass
--fcc-editable-region--

expenses = []
```
