---
id: 65823634a308cb4d6bcc5fa6
title: Paso 21
challengeType: 20
dashedName: step-21
---

# --description--

Las funciones lambda pueden combinarse valiosamente con la función `map()`, que ejecuta una función especificada para cada elemento en una colección de objetos, como una lista:

```py
map(lambda x: x * 2, [1, 2, 3])
```

La función a ejecutar se pasa como el primer argumento, y el iterable se pasa como el segundo argumento.

El resultado del ejemplo anterior sería `[2, 4, 6]`, donde cada elemento en la lista pasada a `map()` ha sido duplicado por la acción de la función lambda.

Modifica tu llamada a `print()` para imprimir el resultado de llamar a `map()` con `test` como primer argumento y `[2, 3, 5, 8]` como segundo argumento. Aún no podrás ver una salida legible. 

# --hints--

Deberías tener `print(map(test, [2, 3, 5, 8]))` en tu código.

```js
({ test: () => assert.match(code, /^print\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(test(3))
--fcc-editable-region--

expenses = []
```
