---
id: 6582267ef39f5047411c1743
title: Paso 14
challengeType: 20
dashedName: step-14
---

# --description--

A continuación, define una función llamada `print_expenses` que recibe un parámetro `expenses`. Esta función se usará más adelante para mostrar cada gasto en tu lista.

Llena el cuerpo de tu nueva función con una sentencia `pass`.

# --hints--

Debes definir una nueva función llamada `print_expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(print_expenses)
  `))
})
```

Tu función `print_expenses` debe tomar un solo parámetro `expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(print_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
