---
id: 66314f2add0a8b57da5d6d91
title: Paso 22
challengeType: 20
dashedName: step-22
---

# --description--

Deberías ver algo como `<map object at 0xd273a8>` impreso en la consola, que es la representación en cadena del objeto map devuelto por `map()`.

Para obtener una salida legible necesitas convertir el objeto map en una lista. Hazlo pasando la llamada `map()` como argumento a la función `list()`.

# --hints--

Debes pasar la llamada `map()` como argumento a la función `list()`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(list(map(test, [2,3,5,8])))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(map(test, [2, 3, 5, 8]))
--fcc-editable-region--

expenses = []
```
