---
id: 66694bfa435cef241f6f9a68
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

Antes de comenzar con el proyecto, vas a aprender sobre listas.

A diferencia de los números y las cadenas, una <dfn>lista</dfn> es un tipo de dato que funciona como un contenedor para otros valores:

```py
example_list = [1, 2, 3, 4]
empty_list = []
```

La lista se caracteriza por los corchetes alrededor de todos los valores y una coma entre los valores, como: `["A", "happy", "list"]`. Si la lista no contiene ningún valor, entonces es una lista vacía: `[]`.

Una lista puede contener diferentes tipos de datos: `[1, "Up", ["Down", "Twice"]]`. Eso incluye todos los tipos de datos posibles. ¡También puede contener otra lista!

Crea una variable llamada `my_list` y asígnale una lista vacía.

# --hints--

Debes tener una variable llamada `my_list`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('my_list'), "variable my_list not found"
    `)
})
```

La variable `my_list` debe tener el valor de una lista vacía.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent('my_list = []'), "variable my_list is not an empty list"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
