---
id: 666951d2e4b27e3a3c72ca5f
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

También puedes acceder a un solo elemento para obtener su valor.

Las listas están indexadas desde cero, al igual que las cadenas. Eso significa que el primer elemento está en el índice `0`, el segundo elemento está en el índice `1` y así sucesivamente.

Para acceder a un elemento usas la notación de corchetes. Por ejemplo, `example_list[1]` está accediendo al elemento en el índice `1`, el segundo elemento, de `example_list`.

Imprime el primer elemento de `my_list`.

# --hints--

Debes imprimir el primer elemento de la lista con `print(my_list[0])`.

```js
({
    test: () => runPython(`
calls =  _Node(_code).find_calls("print")
assert any(call.is_equivalent('print(my_list[0])') for call in calls), "print of my_list[0] not found"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)
--fcc-editable-region--
```
