---
id: 66694eaaeb7d253049fa7883
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Para agregar elementos a una lista, puedes escribirlos entre los corchetes, separando los valores con comas.

Agrega los números `1` y `2` dentro de la lista.

# --hints--

El primer elemento de `my_list` debe ser `1`.

```js
({
    test: () => runPython(`
        assert my_list[0] == 1, "my_list[0] is not 1"
    `)
})
```

El segundo elemento de `my_list` debe ser `2`.

```js
({
    test: () => runPython(`
        assert my_list[1] == 2, "my_list[1] is not 2"
    `)
})
```

`my_list` debería tener dos elementos dentro.

```js
({
    test: () => runPython(`
        assert len(my_list) == 2, "my_list doesn't have two elements"
    `)
})
```

Debes asignar `[1, 2]` a `my_list` directamente.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = []
--fcc-editable-region--
```
