---
id: 6669539c1379793f9cb8917c
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Las listas de Python son mutables, lo que significa que el valor de los elementos de la lista puede cambiarse. Puedes cambiar el valor de un elemento usando la notación de corchetes.

```py
example_list = [4, 5, 6, 7]
example_list[1] = 'oh'
```

Esto hará que `example_list` tenga el valor de `[4, 'oh', 6, 7]`.

Cambia el primer elemento de `my_list` a `0`, luego imprime la lista para verificar el valor.

# --hints--

No cambies la declaración de la lista.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

El primer elemento de `my_list` debe ser 0.

```js
({
    test: () => runPython(`
        assert my_list[0] == 0
    `)
})
```

Debes cambiar el primer elemento usando `my_list[0] = 0`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_stmt("my_list[0] = 0")
    `)
})
```

Debes imprimir la lista.

```js
({
    test: () => runPython(`
        calls = _Node(_code).find_calls('print')
        assert any(call.is_equivalent("print(my_list)") for call in calls)
    `)
})
```

Debes imprimir la lista después de cambiar el primer elemento.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered("my_list[0] = 0", "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

--fcc-editable-region--
```
