---
id: 66694fc4bba24f33ca01fa5b
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Python tiene varios métodos para listas. Como métodos para agregar o eliminar elementos de la lista.

Puedes agregar un elemento al final de una lista usando el método `append()`. Por ejemplo:

```py
example_list = [4, 5, 6]

example_list.append(7)
```

`example_list` ahora es `[4, 5, 6, 7]`.

Intenta usar el método `append()` para agregar `3` a `my_list`. Luego imprime la lista.

# --hints--

El último elemento de `my_list` debe ser `3`.

```js
({
    test: () => runPython(`
        assert my_list[2] == 3, "my_list[2] is not 3"
    `)
})
```

`my_list` debería tener una longitud de 3.

```js
({
    test: () => runPython(`
        assert len(my_list) == 3, "my_list doesn't have three elements"
    `)
})
```

No debes cambiar la declaración de la lista.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

Debes imprimir la lista.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls("print")
    assert any(call.is_equivalent('print(my_list)') for call in calls)
    `)
})
```

Debes imprimir la lista después de usar `append()`.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.append(3)', 'print(my_list)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]
--fcc-editable-region--
```
