---
id: 666ab82c874add3c95145213
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

El método `pop()` se puede usar para eliminar un elemento de una lista. Por defecto, elimina el último elemento de la lista. Puedes pasar un índice como argumento al método, y eliminará el elemento en el índice dado.

```py
fruits_list = ["cherry", "lemon", "tomato", "apple", "orange"]

fruits_list.pop(2)

print(fruits_list) # ["cherry", "lemon", "apple", "orange"]
```

En este caso, `fruits_list.pop(2)` elimina el elemento en el índice `2` de la lista.

Usa `pop()` para eliminar el último elemento de `my_list`, luego imprime `my_list`.

# --hints--

Debes usar `pop()`.

```js
({
    test: () => runPython(`
    assert len(_Node(_code).find_calls('pop')) > 0
    `)
})
```

`my_list` debería ser `[0, 1, 2]`.

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2]
    `)
})
```

Debes imprimir `my_list` después de usar `pop()`.

```js
({
    test: () => runPython(`
    code = _Node(_code)
    assert(
        any(
            code.is_ordered(pop, 'print(my_list)') for pop in ['my_list.pop()', 'my_list.pop(3)']
        )
    )
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

my_list.insert(1, 1)
print(my_list)


--fcc-editable-region--
```
