---
id: 666aae6a5d30a71f1fd7749f
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

El método `insert` puede agregar un elemento en cualquier posición de una lista. El primer argumento es la posición en la que se debe agregar el elemento, y el segundo argumento es el elemento a agregar. Por ejemplo, aquí se muestra cómo agregar un nuevo elemento en la tercera posición de `example_list`:

```py
example_list = [4, 5, 6, 7]

example_list.insert(2, 5.5)

print(example_list) # [4, 5, 5.5, 6, 7]
```

Usando `insert()`, agrega `1` a `my_list` en la posición correcta para que esté contando hacia arriba, luego imprime la lista.

# --hints--

Debes agregar `1` a `my_list` en el índice `1` usando `insert()`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('my_list.insert(1, 1)')
    `)
})
```

`my_list` debería tener 4 elementos.

```js
({
    test: () => runPython(`
    assert len(my_list) == 4
    `)
})
```

`my_list` debería ser igual a `[0, 1, 2, 3]`.

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2, 3]
    `)
})
```

Debes imprimir la lista después de agregar el nuevo elemento.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.insert(1, 1)', "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

--fcc-editable-region--
```
