---
id: 684cd59ec94b5e063bf776f0
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

在那个浅红色背景周围有一个边框会很好。 

要在 Tailwind CSS 中设置 `border-width`，你可以使用 `border-number` 工具类，其中 `number` 将表现你的边框的轻度或厚度。 

下面是一个在所有边应用 `border-width` 为 `4px` 的示例：

```html
<p class="border-4">This is how to set the border-width.</p>
```

在你的开头 `div` 标签内，为所有边设置 `2px` 的 `border-width` 添加正确的实用类。

# --hints--

你的开头 `div` 标签应具有 `border-2` 类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
