---
id: 684cff0bf0a4ca0ad7e42990
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

下一步是将边框的角进行圆角处理。要使用 Tailwind CSS 设置 `border-radius` 属性，你可以使用一个符合以下模式的工具类：`rounded-size`，其中 `size` 表示半径。 

下面是为一个元素应用超大边框半径的示例：

```html
<p class="rounded-xl">This is border radius.</p>
```

在你的开头 `div` 标签内，应用大小为 `md` 的边框半径。 

# --hints--

你的开头 `div` 标签应包含 `rounded-md` 工具类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("rounded-md"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
