---
id: 684cd8bf5445de089af9c837
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

现在是应用边框颜色的时候了。要使用 Tailwind CSS 设置 `border-color` 属性，你可以使用一个具有以下模式的实用类：`border-color-number`，其中 `number` 表现边框颜色的色度。 

下面是为一个元素应用紫色边框颜色的示例：

```html
<p class="border-purple-400">This is border-color.</p>
```

在你的开头 `div` 标签内，应用一个色号为 `300` 的红色边框颜色。 

# --hints--

你的开头 `div` 标签应具有 `border-red-300` 工具类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-red-300"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
