---
id: 684cca09aac0fc035db4099f
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `p` 元素内，添加文本 `Error! Something went wrong. Please try again.`，其中 `Error!` 应该包裹在一个 `strong` 元素内。

# --hints--

你的 `p` 元素内应包含一个 `strong` 元素。

```js
assert.exists(document.querySelector("p strong"));
```

你的 `strong` 元素应包含文本 `Error!`。

```js
assert.equal(document.querySelector("strong")?.textContent, "Error!");
```

你的 `p` 元素应包含以 `Something went wrong. Please try again.` 结尾的正确文本。

```js
const remainingText = "Something went wrong. Please try again."
const pElement = document.querySelector("p");
assert.isTrue(pElement?.textContent.trim().endsWith(remainingText));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      --fcc-editable-region--
      <p>

      </p>
      --fcc-editable-region--
    </div>
  </body>
</html>
```
