---
id: 684cd027744c2a0520ba4d58
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

Tailwind CSS 使用直接应用于标记的实用类。这些类包括字体、颜色、Flexbox 版型、网格等的样式。 

要在 Tailwind CSS 中应用背景颜色，你可以使用 `bg` 前缀，后跟颜色名称和数字。该数字表现颜色的色度。

这是为段落元素应用深蓝色背景的示例。

```html
<p class="bg-blue-900">This has a dark blue background.</p>
```

在你的开头 `div` 标签中，添加一个带有实用类 `bg-red-100` 的类，该类表现为浅红色背景色。

# --hints--

你的开头 `div` 标签应具有一个类属性。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.hasAttribute("class"));
```

你的类属性应包含实用类 `bg-red-100`。

```js
const divEl = document.querySelector("div");
assert.exists(document.querySelector("div.bg-red-100"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
