---
id: 684d09cad4ec900e41db2e74
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

现在，`div` 元素被放置在预览屏幕的左上角。但为元素顶部添加一些空间会更好。要使用 Tailwind CSS 为元素应用 `margin`，你可以使用以下模式：`m-number`，其中每个 `number` 单位代表 0.25rem（或 4 像素），用于应用到四个边。

如果你只想为一侧应用边距，可以使用以下任意模式：`mt-number`（上边距）、`mb-number`（下边距）、`mr-number`（右边距）和 `ml-number`（左边距）。

下面是一个为元素应用 32 像素 `margin-bottom` 的示例：

```html
<p class="mb-8">This is margin.</p>
```

在你的打开的 `div` 标签内，使用正确的工具类来应用 `16px` 的 `margin-top`。

# --hints--

你的开头 `div` 标签应具有 `mt-4` 实用类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mt-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
