---
id: 684d105e488dbd1166a717ae
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

如果你增加预览窗口的大小，你会注意到所有的项仍然在容器中左对齐。如果 flex 项居中会更好。 

要使用 Tailwind CSS 沿容器的主轴定位弹性项，你可以使用以下模式：`justify-value`。`value` 可以是 `start`、`end`、`center`、`between` 等。任何在原生 CSS 中 `justify-content` 使用的有效值都可以在 Tailwind CSS 中使用。 

在你的开头 `div` 标签内，使用正确的实用类来设置 `justify-content` 属性为 `center`。

# --hints--

你的开头 `div` 标签应具有 `justify-center` 工具类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("justify-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
