---
id: 684d17a3b350bb13d3c19c0f
title: 步骤 19
challengeType: 0
dashedName: step-19
---

# --description--

现在是为 `button` 元素应用一些样式的时候了。 

首先为你的 `button` 元素添加一个类属性。接下来，应用正确的实用类以设置红色文本颜色，色调设置为 `700`。

然后，在你的开头 `button` 标签内，使用正确的实用类将 `font-size` 设置为 `2xl`，它表现为 `1.5rem` 或（`24px`）。

# --hints--

你的 `button` 元素应具有一个 class 属性。

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.hasAttribute("class"));
```

你的开头 `button` 标签应包含 `text-red-700` 工具类。

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.classList.contains("text-red-700"));
```

你的开头 `button` 标签应具有 `text-2xl` 实用类。

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl.classList.contains("text-2xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--
      <button aria-label="Dismiss">
      --fcc-editable-region--
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
