---
id: 684d0eeb9c81e5109e52fe8b
title: 步骤15
challengeType: 0
dashedName: step-15
---

# --description--

`div` 元素开始看起来更好了。但现在关闭 `button` 离段落元素有点太近。 

要在 flex 项之间创建间距，你可以使用以下模式：`gap-number`。`number` 表现为 flex 版型中行和列之间的固定间距。

下面是为一个元素应用 `gap` 的示例：

```html
<p class="flex gap-8">This is gap.</p>
```

在你的开头 `div` 标签内，使用正确的实用类将 `gap` 设置为 `4`。

# --hints--

你的开头 `div` 标签应包含 `gap-4` 实用类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("gap-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
