---
id: 684d07cbfa43a20d6a002ed7
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

`div` 元素的样式开始成型了。但如果文本周围有一些间距会更好。 

要使用 Tailwind CSS 为元素应用 `padding`，你可以使用以下模式：`p-number`，其中 `number` 表现你想在四个边上应用的像素。 

这是一个将 `8px` 的 `padding` 应用到元素的示例：

```html
<p class="p-8">This is padding.</p>
```

在你的开头 `div` 标签内，使用正确的实用类来应用 `4px` 的 `padding`。

# --hints--

你的开头 `div` 标签应具有 `p-4` 实用类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
