---
id: 684d009994a3840b98c9aa3a
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

目前，`div` 元素的宽度跨越了整个页面的宽度。要在 Tailwind CSS 中应用宽度，你有几个选项，比如使用 `w-number` 模式来设置固定宽度，或使用 `w-fraction` 模式来设置百分比宽度。  

这是一个将宽度设置为其父容器 `66%` 的示例：

```html
<p class="w-2/3">This is a percentage width.</p>
```

在你的开头 `div` 标签内，应用宽度设置为其父容器的 `50%`。

# --hints--

你的开头 `div` 标签应具有 `w-1/2` 实用类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
