---
id: 684d0e11094c3e0fe1fb98f5
title: 步骤14
challengeType: 0
dashedName: step-14
---

# --description--

现在，段落元素堆叠在关闭 `button` 元素的上方。但如果 `button` 元素能与段落在同一行会更好。这就是 Flexbox 的用武之地。

要将 `div` 元素设置为 Tailwind CSS 中的 Flexbox 版型，你可以使用 `flex` 工具类。

# --hints--

你的开头 `div` 标签应具有 `flex` 实用类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("flex"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
