---
id: 684d160a2370b01310ae0fda
title: 步骤18
challengeType: 0
dashedName: step-18
---

# --description--

新的红色看起来不错，但如果文本更大会更好。要使用 Tailwind CSS 为元素设置 `font-size`，你可以使用以下模式：`text-value`。 

下面是将一个元素的 `font-size` 设置为 `0.875rem` 或 (`14px`) 的示例：

```html
<p class="text-sm">This is for font-size.</p>
```

在你的开头 `p` 元素内，使用正确的实用类将 `font-size` 设置为 `xl`，它表现为 `1.25rem` 或 (`20px`)。

# --hints--

你的开头 `p` 标签应具有 `text-xl` 工具类。

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl.classList.contains("text-xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p class="text-red-700">
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
