---
id: 684d0d369e2b3b0f288d2c20
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

将 `div` 元素水平居中会很不错。要同时为两个方向应用 `margin`，你可以使用以下模式：`mx-number` 或 `my-number`。 

`mx` 设置左右边距，`my` 设置上下边距。

下面是为元素应用顶部和底部边距的示例：

```html
<p class="my-8">This is margin.</p>
```

在你的开头 `div` 标签内，使用正确的实用类将左右外边距设置为 `auto`。

# --hints--

你的开头 `div` 标签应具有 `mx-auto` 实用类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
