---
id: 684d03df2a0ddf0c735865f0
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

现在你看到元素宽度是 `body` 宽度的一半。这对于较大设备效果很好，但对于较小设备效果不佳。

Tailwind CSS 采用移动优先的设计方法。所以如果你想为其他设备创建响应式设计，你可以像这样使用 Tailwind 的响应式变体，比如 `md` 和 `lg`：

```html
<p class="lg:w-2/3">This is a percentage width.</p>
```

以下示例将在设备尺寸为 `1024px` 及更高时，将宽度设置为父元素的 `66%`。 

在你的开头 `div` 标签内，修改你的 `w-1/2`，使其仅应用于 `768px` 及以上的设备尺寸。你可以为这个类添加 `md` 变体前缀。 

# --hints--

你的开头 `div` 标签应具有 `md:w-1/2` 实用类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:w-1/2"));
```

你的开头 `div` 标签中不应再有 `w-1/2` 类。

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
