---
id: 66f26c32ec6f90df01a44f60
title: 设计一个联系表单
challengeType: 25
dashedName: design-a-contact-form
demoType: onClick
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该有一个类为 `form-container` 的 `div`，用于放置表单。

2. 你在 `.form-container` 元素中应该有一个 `form` 元素。

3. 在 `form` 元素中，应包含以下元素和输入字段：

   - 一个包含一些文本的 `h2` 元素。
   - 一个姓名文本输入字段，`type` 设置为 `text`，并包含 `id`、`name`、`required` 属性。
   - 一个电子邮件输入字段，`type` 设置为 `email`，并包含 `id`、`name`、`required` 属性。
   - 一个信息文本区域，包含 `id`、`name`、`required` 属性。
   - 一个 `button` 元素，文本为 `"Submit"`，`type` 属性为 `submit`。

4. 每个 `input` 和 `textarea` 元素都应该有其相应的 `label` 元素，将 `for` 属性值设置为元素的 `id` 值。

5. `.form-container` 元素应具有背景色，并且 `border-radius`、`padding` 和 `width` 都有值。

6. `label` 元素应有外边距和字体颜色。

7. `input` 和 `textarea` 元素应具有以下值：

   - `width`。
   - `padding`。
   - 底部外边距。

8. `button` 元素应具有：

   - 背景色
   - 非默认的字体大小。

9.  按钮应具有改变背景颜色的悬停效果。

**注意：** 确保在 HTML 中链接样式表以应用 CSS。

# --hints--

你应该有一个类为 `form-container` 的 `div`。

```js
const el = document.querySelector('div.form-container');
assert.isNotNull(el);
```

你在 `.form-container` 元素中应该有一个 `form` 元素。

```js
const el = document.querySelector('.form-container form');
assert.isNotNull(el);
```

在 `form`中，你应该有一个包含一些文本的 `h2` 元素。

```js
const el = document.querySelector('.form-container form h2');
assert.isNotEmpty(el.innerText);
// assert.isAtLeast(document.querySelector('h1').innerText.length, 1);

```

在 `form`中，你应该有一个类型设置为 `text` 的文本输入字段。

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert.isNotNull(el);
```

你的 `type=text` 输入字段应该有一个 `id`、一个 `name` 和一个 `required` 属性。

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert.exists(el);
assert.isTrue(el.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

在 `form`中，你应该有一个电子邮件输入字段，`type` 设置为 `email`。

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert.isNotNull(el);
```

你的 `type=email` 输入字段应该有一个 `id`、一个 `name` 和一个 `required` 属性。

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert.isTrue(el?.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

在 `form`中，你应该有一个 `textarea`，包含一个 `id`、一个 `name` 和一个 `required` 属性。

```js
const el = document.querySelector('.form-container form textarea');
assert.isNotNull(el);
assert.isNotNull(document.querySelector('.form-container form textarea[required]'));
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

在 `form`中，你应该有一个 `button` 元素，`type` 属性值为 `submit`。

```js
const el = document.querySelector('.form-container form button');
assert.strictEqual(el?.getAttribute('type'), 'submit');
```

你的按钮文本应为 `"Submit"` 。

```js
const el = document.querySelector('.form-container form button');
assert.strictEqual(document.querySelector('.form-container form button').textContent, 'Submit');
```

你应该有一个 `label` 元素与 `#name` 元素关联，其 `for` 属性与相应 `input` 的 `id` 值相匹配。

```js
const nameInput = document.querySelector('.form-container form input[type="text"]')?.getAttribute('id');

const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

// check if id has a value
assert.isNotEmpty(nameInput);
let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === nameInput) {
        found = true;
        break;
    }
}

assert.isTrue(found);
```

你应该有一个 `label` 元素与 `#email` 元素关联，其 `for` 属性与相应 `input` 的 `id` 值相匹配。

```js
const emailInput = document.querySelector('.form-container form input[type="email"]')?.getAttribute('id');
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);
// check if id has a value
assert.isNotEmpty(emailInput);
let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === emailInput) {
        found = true;
        break;
    }
}

assert.isTrue(found);
```

你应该有一个 `label` 元素与 `#message` 元素关联，其 `for` 属性与相应 `textarea` 的 `id` 值相匹配。

```js
const messageInput = document.querySelector('.form-container form textarea')?.getAttribute('id');
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

// check if id has a value

assert.isNotEmpty(messageInput);

assert.exists(document.querySelector(`label[for="${messageInput}"]`));
```

标签不应为空。

```js
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

for (let label of labels) {
    const labelText = label.textContent.trim();
    assert.isNotEmpty(labelText);
}

```

`.form-container` 元素应具有背景颜色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.backgroundColor);
```

`.form-container` 元素应具有 `border-radius`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.borderRadius);
```

`.form-container` 元素应具有内边距。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.padding);
```

`.form-container` 元素应具有宽度。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.width);
```

标签应有外边距和字体颜色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('label')?.margin);
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('label')?.color);
```

你的 `input` 元素应有宽度、内边距和底部外边距。

```js
// Helper function to extract properties from multiple selectors
function getStyleProperties(selectors, properties) {
  const result = {};
  properties.forEach((property) => {
    result[property] = selectors.map((selector) => selector?.[property]).find(Boolean);
  });
  return result;
}

// Initializing selectors
const selectors = [
  new __helpers.CSSHelp(document).getStyle('input'),
  new __helpers.CSSHelp(document).getStyle('input, textarea'),
  new __helpers.CSSHelp(document).getStyle('textarea, input')
];

// Extract width, padding, and margin properties
const { width } = getStyleProperties(selectors, ['width']);
assert.isNotEmpty(width);

const {
  padding,
  paddingTop,
  paddingBottom,
  paddingLeft,
  paddingRight
} = getStyleProperties(selectors, ['padding', 'paddingTop', 'paddingBottom', 'paddingLeft', 'paddingRight']);

// Assert padding properties
assert.isOk(
  padding ||
  (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

const { marginBottom } = getStyleProperties(selectors, ['marginBottom']);
assert.isNotEmpty(marginBottom);

```

你的 `textarea` 元素应有宽度、内边距和底部外边距。

```js
// Helper function to extract properties from multiple selectors
function getStyleProperties(selectors, properties) {
  const result = {};
  properties.forEach((property) => {
    result[property] = selectors.map((selector) => selector?.[property]).find(Boolean);
  });
  return result;
}

// Initializing selectors
const selectors = [
  new __helpers.CSSHelp(document).getStyle('textarea'),
  new __helpers.CSSHelp(document).getStyle('input, textarea'),
  new __helpers.CSSHelp(document).getStyle('textarea, input')
];

// Extract width, padding, and margin properties
const { width } = getStyleProperties(selectors, ['width']);
assert.isNotEmpty(width);

const {
  padding,
  paddingTop,
  paddingBottom,
  paddingLeft,
  paddingRight
} = getStyleProperties(selectors, ['padding', 'paddingTop', 'paddingBottom', 'paddingLeft', 'paddingRight']);

// Assert padding properties
assert.isOk(
  padding ||
  (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

const { marginBottom } = getStyleProperties(selectors, ['marginBottom']);
assert.isNotEmpty(marginBottom);
```

你的 `button` 元素应有背景颜色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button')?.backgroundColor);
```

你的 `button` 元素应有字体大小。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button')?.fontSize);
```

你的 `button` 元素应有改变背景颜色的悬停效果。

```js
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
const btn_hover = gs('button:hover')?.backgroundColor;
assert.isNotEmpty(btn_hover);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Contact Form</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Styled Form Project</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="form-container">
        <form>
            <h2>Contact Us</h2>
            <label for="name">Name:</label>
            <input type="text" id="name" name="name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" required>

            <label for="message">Message:</label>
            <textarea id="message" name="message" required></textarea>

            <button type="submit">Submit</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background: linear-gradient(45deg, #83a4d4, #b6fbff);
    font-family: Arial, sans-serif;
}

.form-container {
    background-color: white;
    border-radius: 10px;
    padding: 20px;
    width: 300px;
    text-align: center;
}

form h2 {
    margin-bottom: 20px;
    color: #333;
}

label {
    display: block;
    margin: 10px 0 5px;
    color: #555;
}

input{
    width: 100%;
    padding: 10px;
    margin-bottom: 10px;
    border: 1px solid #ddd;
    border-radius: 5px;
    box-sizing: border-box;
    resize: vertical;
}

textarea {
    width: 100%;
    padding: 10px;
    margin-bottom: 10px;
    border: 1px solid #ddd;
    border-radius: 5px;
    box-sizing: border-box;
    resize: vertical;
}

button {
    background-color: #4CAF50;
    color: white;
    padding: 10px 15px;
    border: none;
    border-radius: 5px;
    cursor: pointer;
    font-size: 1em;
}

button:hover {
    background-color: #45a049;
}
```
