---
id: 680b3ef395479b0e449ecb6e
title: تنفيذ خوارزمية الترتيب بالاختيار
challengeType: 27
dashedName: implement-selection-sort-algorithm
---

# --description--

الترتيب بالاختيار هو خوارزمية ترتيب شائعة أخرى تُدرس في معظم دورات علوم الكمبيوتر.

تعمل هذه الخوارزمية عن طريق إيجاد أصغر عنصر بشكل متكرر من الجزء غير المرتب من القائمة وتبديله مع أول عنصر غير مرتب. تبدأ باختيار القيمة الأدنى في القائمة بأكملها وتبديلها مع العنصر الأول. ثم تنتقل إلى الموضع الثاني، وتجد أصغر قيمة في العناصر غير المرتبة المتبقية، وتبدلها مع العنصر الثاني. تستمر هذه العملية، متقدمة عبر القائمة عنصرًا تلو الآخر، حتى يتم ترتيب القائمة بالكامل.

ينتج عن الترتيب بالاختيار تعقيد زمني تربيعي في أفضل وأوسط وأسوأ الحالات. سيكون تعقيد الذاكرة ثابتًا `O(1)` لأن الترتيب يتم في مكانه وتُستخدم كمية ثابتة من الذاكرة بغض النظر عن حجم القائمة.

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تجتاز لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تعرف دالة باسم `selection_sort`.
1. يجب أن تحتوي دالة `selection_sort` على معلمة واحدة تمثل قائمة العناصر.
1. يجب أن تأخذ دالة `selection_sort` قائمة وترتب العناصر في مكانها من الأصغر إلى الأكبر.
1. يجب أن تعدل دالة `selection_sort` القائمة المُدخلة في مكانها، وتُعيدها بعد ترتيبها.
1. يجب أن تتبع دالة `selection_sort` خوارزمية الترتيب بالاختيار، بتبديل أصغر عنصر من الجزء غير المرتب من القائمة مع أول عنصر غير مرتب.
1. يجب ألا تقوم دالة `selection_sort` بتبديلات غير ضرورية عندما يكون أصغر عنصر في الموضع الصحيح بالفعل.
1. يجب ألا تستخدم دالة `selection_sort` إما الطريقة المدمجة `sort()` أو الدالة `sorted()`.

# --hints--

يجب أن تعرف دالة باسم `selection_sort`.

```js
({test: () => {
    runPython(`
assert _Node(_code).has_function('selection_sort')
    `)
}})
```

يجب أن تحتوي دالة `selection_sort` على معلمة واحدة.

```js
(
    {test: () => {
        runPython(`
            import inspect
            sig = inspect.signature(selection_sort)
            assert len(sig.parameters) == 1
        `)
    }}
)
```

يجب ألا تستورد أي وحدة أو تستخدم طرق الترتيب المدمجة في شفرتك.

```js
(
    {
        test: () => runPython(`
            assert len(_Node(_code).find_imports()) == 0
            assert not _Node(_code).block_has_call("sort")
            assert not _Node(_code).block_has_call("sorted")
        `)
    }
)
```

يجب أن تُعيد دالة `selection_sort` نفس القائمة التي تم إدخالها.

```js
(
    {
        test: () => runPython(`
            input_list = [32, 0, 88, 1, 203]
            assert selection_sort(input_list) is input_list
        `)
    }
)
```

يجب أن تعدل دالة `selection_sort` القائمة المُدخلة في مكانها. لا تستخدم أي طريقة لإضافة أو إزالة عناصر من القائمة.

```js
(
    {
        test: () => runPython(`
            from collections import UserList

            class CustomList(UserList):
                def __init__(self, *args, **kwargs):
                    self._record = []
                    super().__init__(*args, **kwargs)

                def __setitem__(self, key, value):
                    self._record.append(f'Setting item {key} to {value}. Previous value: {self[key]}')
                    super().__setitem__(key, value)

                def append(self, *args, **kwargs):
                    self._record.append(f'Appending to list: {args}, {kwargs}')
                    super().append(*args, **kwargs)

                def insert(self, *args, **kwargs):
                    self._record.append(f'Inserting to list: {args}, {kwargs}')
                    super().insert(*args, **kwargs)

                def pop(self, *args, **kwargs):
                    self._record.append(f'Popping from list: {args}, {kwargs}')
                    super().pop(*args, **kwargs)

                def remove(self, *args, **kwargs):
                    self._record.append(f'Removing from list: {args}, {kwargs}')
                    super().remove(*args, **kwargs)


            list_to_sort = CustomList([32, 0, 88, 1, 203])

            selection_sort(list_to_sort)
            assert all(
                'Setting item' in record
                for record in list_to_sort._record
            )
        `)
    }
)
```

يجب أن تتبع دالة `selection_sort` خوارزمية الترتيب بالاختيار، بتبديل القيمة الأدنى في الجزء غير المرتب من القائمة مع أول عنصر غير مرتب. تجنب التبديلات غير الضرورية عندما تكون القيمة الأدنى في الموضع الصحيح بالفعل.

```js
(
    {
        test: () => runPython(`
            from collections import UserList

            class CustomList(UserList):
                def __init__(self, *args, **kwargs):
                    self._record = []
                    super().__init__(*args, **kwargs)

                def __setitem__(self, key, value):
                    self._record.append(f'Setting item {key} to {value}. Previous value: {self[key]}')
                    super().__setitem__(key, value)


            list_to_sort = CustomList([33, 1, 89, 2, 67, 245])

            swap_pairs = [
                ((0, 1, 33), (1, 33, 1)),
                ((1, 2, 33), (3, 33, 2)),
                ((2, 33, 89), (3, 89, 33)),
                ((3, 67, 89), (4, 89, 67)),
            ]

            selection_sort(list_to_sort)
            record = list_to_sort._record

            actual_pairs = [
                (record[offset * 2], record[offset * 2 + 1])
                for offset, _ in enumerate(record[::2])
            ]
            assert len(swap_pairs) == len(actual_pairs)

            for expected_pair, actual_pair in zip(swap_pairs, actual_pairs):
                assert (
                    set(f'Setting item {key} to {value}. Previous value: {old}' for key, value, old in expected_pair)
                    == set(actual_pair)
                )
        `)
    }
)
```

`selection_sort([33, 1, 89, 2, 67, 245])` يجب أن تُعيد `[1, 2, 33, 67, 89, 245]`.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([33, 1, 89, 2, 67, 245]) == [1, 2, 33, 67, 89, 245]
            
        `)
    }
)
```

`selection_sort([5, 16, 99, 12, 567, 23, 15, 72, 3])` يجب أن تُعيد `[3, 5, 12, 15, 16, 23, 72, 99, 567]`.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([5, 16, 99, 12, 567, 23, 15, 72, 3]) == [3, 5, 12, 15, 16, 23, 72, 99, 567]
            
        `)
    }
)
```

`selection_sort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` يجب أن تُعيد `[1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]`.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]) == [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
            
        `)
    }
)
```

يجب أن ترتب دالة `selection_sort` بشكل صحيح أي قائمة أرقام.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([42, 17, 93, 8, 61, 29]) == [8, 17, 29, 42, 61, 93]
            assert selection_sort([11, 4, 78, 23, 55, 198, 65, 90, 2]) == [2, 4, 11, 23, 55, 65, 78, 90, 198]
            assert selection_sort([9, 27, 3, 7, 101, 66, 34, 52, 87, 42, 12, 29]) == [3, 7, 9, 12, 27, 29, 34, 42, 52, 66, 87, 101]
            assert selection_sort([5, 14, 33, 77, 2, 18, 92, 1, 100, 45, 73, 64, 28, 56]) == [1, 2, 5, 14, 18, 28, 33, 45, 56, 64, 73, 77, 92, 100]
            
        `)
    }
)     
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def selection_sort(nums):
    for i, _ in enumerate(nums):
        min_index = i

        for j in range(i + 1, len(nums)):
            if nums[j] < nums[min_index]:
                min_index = j

        if min_index != i:
            nums[i], nums[min_index] = nums[min_index], nums[i]

    return nums

```
