---
id: 66e2e5d890ac8a54c9d703bd
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

É hora de contar as consoantes. Crie uma função `getConsonantCount` com um parâmetro `sentence`.

Dentro da função, use um loop para contar o número de consoantes na `sentence` que será passada para a função quando ela for chamada. Uma consoante é qualquer letra que não seja um dos seguintes caracteres: `"aeiou"`. 

Sua função `getConsonantCount` deve retornar um número.

# --hints--

Você deve criar uma função `getConsonantCount`.

```js
assert.isFunction(getConsonantCount);
```

Sua função `getConsonantCount` deve ter um parâmetro `sentence`.

```js
assert.match(getConsonantCount.toString(), /sentence/);
```

Sua função `getConsonantCount` deve retornar um número.

```js
assert.isNumber(getConsonantCount("Coding is fun"))
```

Quando a frase for `"Coding is fun"`, a função `getConsonantCount` deve retornar `7`.

```js
assert.strictEqual(getConsonantCount("Coding is fun"), 7);
```

Quando a frase for `"Hello, World!"`, a função `getConsonantCount` deve retornar `7`.

```js
assert.strictEqual(getConsonantCount("Hello, World!"), 7);
```

A contagem de consoantes deve ser case-insensitive.

```js
assert.strictEqual(getConsonantCount("Apples are tasty fruits"), 13);
assert.strictEqual(getConsonantCount("freeCodeCamp is awesome"), 11);
```

Sua função `getConsonantCount` deve retornar a contagem correta de consoantes para qualquer frase.

```js
assert.strictEqual(getConsonantCount("I went to the store"), 9);
assert.strictEqual(getConsonantCount("The quick brown fox jumps over the lazy dog"), 24);
assert.strictEqual(getConsonantCount("The cat in the hat"), 9);
assert.strictEqual(getConsonantCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
