---
id: 66e2e99caa0cb3570f4822b4
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Agora é hora de testar sua função `getConsonantCount`.

Crie uma variável `consonantCount` e atribua a ela o resultado da chamada da função `getConsonantCount` com o argumento `"Coding is fun"`

Depois disso, registre o seguinte no console: `"Consonant Count: [Consonant count goes here]"`. Substitua `[Consonant count goes here]` pelo nome real da variável.  You can choose to use template strings or string concatenation with the `+` operator here.

# --hints--

Você deve criar uma variável `consonantCount`.

```js
assert.isNotNull(consonantCount)
```

Sua variável `consonantCount` deve ser definida com o resultado de `getConsonantCount("Coding is fun")`.

```js
assert.equal(consonantCount, getConsonantCount("Coding is fun"));
```

Você deve registrar a variável `consonantCount` no console com o parâmetro `"Consonant Count: ${consonantCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Consonant\s+Count:\s+('|"|`)\s+\+\s+consonantCount|`Consonant\s+Count:\s+\${consonantCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
