---
id: 66e2edb23013ca596f8d9568
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Agora é hora de testar sua função `getPunctuationCount`.

Crie uma variável `punctuationCount` e atribua a ela o resultado da chamada da função `getPunctuationCount` com o argumento `"WHAT?!?!?!?!?"`

Depois disso, registre o seguinte no console: `"Punctuation Count: [Punctuation count goes here]"`. Substitua `[Punctuation count goes here]` pelo nome real da variável. Você pode optar por usar template strings ou concatenação de strings com o operador `+` aqui.

# --hints--

Você deve criar uma variável `punctuationCount`.

```js
assert.isNotNull(punctuationCount)
```

Sua variável `punctuationCount` deve ser definida com o resultado de `getPunctuationCount("WHAT?!?!?!?!?")`.

```js
assert.equal(punctuationCount, getPunctuationCount("WHAT?!?!?!?!?"));
```

Você deve registrar o seguinte no console: `"Punctuation Count: [Punctuation count goes here]"`. Replace  `[Punctuation count goes here]` with the actual variable name.Certifique-se de usar a sintaxe correta de concatenação de strings aqui.

```js
assert.match(code, /console\.log\((?:('|"|`)Punctuation\s+Count:\s+('|"|`)\s+\+\s+punctuationCount|`Punctuation\s+Count:\s+\${punctuationCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```

