---
id: 66e2df8c5676fe51cead8ec8
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Agora é hora de testar sua função `getVowelCount`.

Crie uma variável `vowelCount` e atribua a ela o resultado da chamada da função `getVowelCount` com o argumento `"Apples are tasty fruits"`

Depois disso, registre o seguinte no console: `"Vowel Count: [vowel count goes here]"`.Substitua `[vowel count goes here]` pelo nome real da variável. You can choose to use template strings or string concatenation with the `+` operator here.

# --hints--

Você deve criar uma variável `vowelCount`.

```js
assert.isNotNull(vowelCount)
```

Sua variável `vowelCount` deve ser definida com o resultado de `getVowelCount("Apples are tasty fruits")`.

```js
assert.equal(vowelCount, getVowelCount("Apples are tasty fruits"));
```

Você deve registrar a variável `vowelCount` no console com o parâmetro `"Vowel Count: ${vowelCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Vowel\s+Count:\s+('|"|`)\s+\+\s+vowelCount|`Vowel\s+Count:\s+\${vowelCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
