---
id: 66e2d680e129e1423116a541
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Neste workshop, você irá construir um analisador de sentenças que receberá uma sentença e obterá a contagem do número de palavras, vogais, consoantes e sinais de pontuação.

Para começar, crie uma função `getVowelCount` com um parâmetro chamado `sentence`.Sua função deve retornar o número total de vogais em uma frase.

# --hints--

Você deve criar uma função `getVowelCount`.

```js
assert.isFunction(getVowelCount);
```

Sua função `getVowelCount` deve ter um parâmetro `sentence`.

```js
assert.match(getVowelCount.toString(), /sentence/);
```

Sua função `getVowelCount` deve retornar um número.

```js
assert.isNumber(getVowelCount("Apples are tasty fruits"))
```

Quando a frase é `"Apples are tasty fruits"`, a função `getVowelCount` deve retornar `7`.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
```

Quando a sentença for `"Hello, World!"`, a função deve retornar `3`.

```js
assert.strictEqual(getVowelCount("Hello, World!"), 3);
```

A contagem de vogais deve ser case-insensitive.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
assert.strictEqual(getVowelCount("freeCodeCamp is awesome"), 10);
```

Sua função `getVowelCount` deve retornar a contagem correta de vogais para qualquer frase.

```js
assert.strictEqual(getVowelCount("I went to the store"), 6);
assert.strictEqual(getVowelCount("The quick brown fox jumps over the lazy dog"), 11);
assert.strictEqual(getVowelCount("The cat in the hat"), 5);
assert.strictEqual(getVowelCount(""), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
