---
id: 66e2eab8a5638f57b637b7cc
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Você deve contar o número de pontuações agora.

Crie uma função `getPunctuationCount` com um parâmetro `sentence`.

Dentro da função, crie um loop para contar o número de caracteres de pontuação na `sentence` que será passada para a função quando ela for chamada. Para nossos propósitos, um caractere de pontuação é qualquer caractere que não seja um espaço (`" "`) ou uma letra.

Sua função `getPunctuationCount` deve retornar um número.

# --hints--

Você deve criar uma função `getPunctuationCount`.

```js
assert.isFunction(getPunctuationCount);
```

Sua função `getPunctuationCount` deve ter um parâmetro `sentence`.

```js
assert.match(getPunctuationCount.toString(), /sentence/);
```

Sua função `getPunctuationCount` deve retornar um número.

```js
assert.isNumber(getPunctuationCount("Coding is fun!"))
```

Quando a frase for `"What's going on here?"`, a função `getPunctuationCount` deve retornar `2`.

```js
assert.strictEqual(getPunctuationCount("What's going on here?"), 2);
```

Quando a frase for `"What????!"`, a função `getPunctuationCount` deve retornar `5`.

```js
assert.strictEqual(getPunctuationCount("What????!"), 5);
```

Sua função `getPunctuationCount` deve retornar a contagem correta de pontuação para qualquer frase.

```js
assert.strictEqual(getPunctuationCount("Be quick, sign up! freeCodeCamp awaits, friend!!!"), 6);
assert.strictEqual(getPunctuationCount("Guess what? freeCodeCamp is launching a new cert soon!"), 2);
assert.strictEqual(getPunctuationCount("freeCodeCamp, again? It's incredible!"), 4);
assert.strictEqual(getPunctuationCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
