---
id: 66e2f06b191e305c00574e4d
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Finalmente, conte o número de palavras criando uma função `getWordCount` com um parâmetro `sentence`. A função deve retornar o número total de palavras na frase passada quando for chamada.

# --hints--

Você deve criar uma função `getWordCount`.

```js
assert.isFunction(getWordCount);
```

Sua função `getWordCount` deve ter um parâmetro `sentence`.

```js
assert.match(getWordCount.toString(), /sentence/);
```

Sua função `getWordCount` deve retornar um número.

```js
assert.isNumber(getWordCount("I love freeCodeCamp"))
```

Quando a frase é `"When are you gonna start learning to code?"`, a função `getWordCount` deve retornar `8`.

```js
assert.strictEqual(getWordCount("When are you gonna start learning to code?"), 8);
```

Quando a frase for `"What's going on?"`, a função `getWordCount` deve retornar `3`.

```js
assert.strictEqual(getWordCount("What's going on?"), 3);
```

A contagem de palavras deve ser case-insensitive.

```js
assert.strictEqual(getWordCount("freeCodeCamp offers free coding tutorials online"), 6);
assert.strictEqual(getWordCount("You can learn HTML, CSS, JavaScript, and more on freeCodeCamp"), 10);
```

Sua função `getWordCount` deve retornar a contagem correta de palavras para qualquer frase.

```js
assert.strictEqual(getWordCount("freeCodeCamp has a great community of kind people"), 8);
assert.strictEqual(getWordCount("The freeCodeCamp curriculum is constantly updated"), 6);
assert.strictEqual(getWordCount("freeCodeCamp teaches both frontend and backend development"), 7);
```

Sua função `getWordCount` deve retornar a contagem correta de palavras para uma string vazia ou uma string contendo apenas espaços.

```js
assert.strictEqual(getWordCount(""), 0);
assert.strictEqual(getWordCount("   "), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
