---
id: 69b279bae0f20eda63163cf8
title: Paso 14
challengeType: 1
dashedName: step-14
---

# --description--

Las sentencias `if` y `else` que creaste en los pasos anteriores asignarán elementos al arreglo ordenado.

Cada elemento asignado al arreglo ordenado ocupa un índice en la lista. Por eso tienes que pasar al siguiente índice en el arreglo ordenado después de cada asignación.

Debajo del bloque `if`/`else`, pero aún dentro del ciclo `while`, incrementa `sortedIndex` en `1`.

# --hints--

Debes incrementar el valor actual de `sortedIndex` en `1` después de tu cláusula `else`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
