---
id: 69b279bae0f20eda63163cf3
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

Dentro del ciclo, comienza a comparar los elementos actuales de cada mitad para que puedas combinar primero el valor más pequeño.

Agrega un bloque `if` que verifique si el elemento actual en `leftPart` es menor que el elemento actual en `rightPart`. 
# --hints--

La condición `if` debe verificar si `leftPart[leftArrayIndex]` es menor que `rightPart[rightArrayIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*<\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
