---
id: 69b279bae0f20eda63163cf6
title: Paso 12
challengeType: 1
dashedName: step-12
---

# --description--

En un paso anterior, asignaste el elemento en el `leftPart` a la lista `array` porque era más pequeño. Pero esto no siempre será así. En algunos casos de comparación, el elemento de la derecha podría ser más pequeño.

Crea una cláusula `else` para ejecutar cuando el elemento en `leftPart` no sea menor que el elemento en `rightPart`.

Dentro del bloque `else`, asigna `rightPart[rightArrayIndex]` a `array[sortedIndex]`.

# --hints--

Debes crear un bloque `else` que contenga `array[sortedIndex] = rightPart[rightArrayIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /else\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
}
```
