---
id: 69b278519ef46b2fd8bf1546
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Ya tienes un valor `middlePoint`. Úsalo para dividir el arreglo en dos mitades y deja la porción izquierda lista para ordenar.

Usa la sintaxis de slice para extraer la mitad izquierda de `array` y asígnala a una variable llamada `leftPart`.

# --hints--

Debes declarar una variable `leftPart` dentro de la función `mergeSort`.

```js
const explorer = await __helpers.Explorer(code);
const { leftPart } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  leftPart,
  'You should declare a leftPart variable inside mergeSort'
);
```

Debes asignar `array.slice(0, middlePoint)` a `leftPart`.

```js
const explorer = await __helpers.Explorer(code);
const { leftPart } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  leftPart,
  'You should declare a leftPart variable inside mergeSort'
);
assert.isTrue(
  leftPart.value.matches('array.slice(0, middlePoint)'),
  'You should assign array.slice(0, middlePoint) to leftPart'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
