---
id: 69b279bae0f20eda63163cfe
title: Paso 20
challengeType: 1
dashedName: step-20
---

# --description--

Antes de probar la función `mergeSort()`, necesitas crear un caso base que detenga la ejecución de la función cuando la longitud de `array` sea menor o igual a `1`.

Este caso base detendrá la llamada recursiva. Sin él, la operación de ordenamiento por mezcla continuaría ejecutándose incluso cuando la lista ya esté ordenada o no tenga elementos.

Justo después de la declaración de la función, crea una sentencia `if` que verifique si `array.length` es menor o igual a 1. Dentro del cuerpo de la sentencia `if`, añade una sentencia `return` para detener la ejecución de la función.

# --hints--

Debes crear una sentencia `if` con la condición `array.length <= 1`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{/
);
```

Debes tener una sentencia `return` dentro del cuerpo de la sentencia `if`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{[\s\S]*return\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }
  
  while (rightArrayIndex < rightPart.length) {
    array[sortedIndex] = rightPart[rightArrayIndex];
    rightArrayIndex += 1;
    sortedIndex += 1;
  }
}
```
