---
id: 69b279bae0f20eda63163cef
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Ahora que has dividido la lista `array` en dos listas separadas, seguirás dividiendo cada lista hasta que cada elemento esté solo en su propia lista. Una lista con un solo número siempre está ordenada.

Para hacerlo, llama recursivamente a `mergeSort` dentro de tu función y pasa `leftPart` como argumento en la llamada.

# --hints--

Debes llamar a `mergeSort(leftPart)` dentro de la misma función `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*leftPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
