---
id: 69b279bae0f20eda63163cf2
title: Paso 8
challengeType: 1
dashedName: step-8
---

# --description--

A continuación, necesitas crear un bucle que continúe mientras haya elementos restantes tanto en `leftPart` como en `rightPart`.

Para ello, crea un bucle `while` con dos condiciones: una que verifique si `leftArrayIndex` es menor que la longitud de `leftPart` y otra condición que verifique si `rightArrayIndex` es menor que la longitud de `rightPart`.

# --hints--

Debes crear un bucle `while` que compruebe si `leftArrayIndex` es menor que la longitud de `leftPart` y si `rightArrayIndex` es menor que la longitud de `rightPart`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
