---
id: 69b2713ff6c21fed2d3ad88b
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

El algoritmo de ordenamiento por mezcla es un algoritmo que utiliza el principio de dividir y conquistar para ordenar colecciones de datos. Es decir, "divide" una colección en sub-partes más pequeñas y "conquista" las sub-partes ordenándolas de forma independiente, luego combina las sub-partes ordenadas.

Comienza creando una función llamada `mergeSort` que reciba el parámetro `array`.

# --hints--

Debes declarar una función llamada `mergeSort` con un solo parámetro `array`.

```js
assert.isFunction(mergeSort);
const regex = __helpers.functionRegex('mergeSort', ['array']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
