---
id: 69b279bae0f20eda63163cf0
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

Al final del cuerpo de tu función, llama de nuevo a la función `mergeSort()`. Esta vez, pasa `rightPart` como argumento en la llamada a la función.

# --hints--

Debes llamar a `mergeSort()` con el argumento `rightPart` al final del cuerpo de tu función.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*rightPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
