---
id: 69b2762b17d23fbee40a6c00
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

El algoritmo de ordenamiento por mezcla realiza tres acciones:

- Dividir una secuencia desordenada de elementos en subpartes
- Ordenar los elementos en las subpartes
- Combinar las subpartes ordenadas

Lo anterior ocurre de forma recursiva hasta que las subpartes se combinan en la secuencia completa ordenada. Comencemos dividiendo la secuencia.

Primero, dentro de la función `mergeSort`, declara una variable `middlePoint` usando `const` y asígnale `Math.floor(array.length / 2)` para que puedas dividir la secuencia a la mitad antes de llamar recursivamente.

# --hints--

Debes declarar una variable `middlePoint` y asignarle `Math.floor(array.length / 2)`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(const|let|var)\s+middlePoint\s*=\s*Math\.floor\(\s*array\.length\s*\/\s*2\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
