---
id: 69b2713ff6c21fed2d3ad88b
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

Merge sort es un algoritmo de ordenamiento que utiliza el principio de dividir y conquistar para ordenar colecciones de datos. Es decir, "divide" una colección en subpartes más pequeñas y "conquista" las subpartes ordenándolas de forma independiente, luego combina las subpartes ordenadas.

Comienza creando una función llamada `mergeSort` que reciba el parámetro `array`.

# --hints--

Debes declarar una función llamada `mergeSort` con un solo parámetro `array`.

```js
assert.isFunction(mergeSort);
const params = __helpers.getFunctionParams(mergeSort.toString());
assert.lengthOf(params, 1);
assert.deepEqual(params[0].name, "array");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
