---
id: 69b278d82f0e514f32acb724
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Después de cortar la mitad izquierda, todavía necesitas la parte derecha para completar el paso de combinación.

Usa la sintaxis de slice para extraer la mitad derecha de `array` y asignarla a una variable llamada `rightPart`.

# --hints--

Deberías tener una variable `rightPart` dentro de la función `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightPart\s*=/
);
```

Debes usar la sintaxis de slice para extraer la mitad derecha del arreglo y asignarla a la variable `rightPart`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /rightPart\s*=\s*array\.slice\(\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
