---
id: 69b278d82f0e514f32acb724
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Después de cortar la mitad izquierda, todavía necesitas la parte derecha para completar el paso de combinación.

Usa la sintaxis de slice para extraer la mitad derecha de `array` y asignarla a una variable llamada `rightPart`.

# --hints--

Deberías tener una variable `rightPart` dentro de la función `mergeSort`.

```js
const explorer = await __helpers.Explorer(code);
const { rightPart } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  rightPart,
  'You should declare a rightPart variable inside mergeSort'
);
```

Debes usar la sintaxis de slice para extraer la mitad derecha del arreglo y asignarla a la variable `rightPart`.

```js
const explorer = await __helpers.Explorer(code);
const { rightPart } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  rightPart,
  'You should declare a rightPart variable inside mergeSort'
);
assert.isTrue(
  rightPart.value.matches('array.slice(middlePoint)'),
  'You should assign array.slice(middlePoint) to rightPart'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
