---
id: 69b2713ff6c21fed2d3ad88b
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

El algoritmo de ordenamiento por mezcla es un algoritmo que utiliza el principio de dividir y conquistar para ordenar colecciones de datos. Es decir, "divide" una colección en sub-partes más pequeñas y "conquista" las sub-partes ordenándolas de forma independiente, luego combina las sub-partes ordenadas.

Comienza creando una función llamada `mergeSort` que reciba el parámetro `array`.

# --hints--

Debes declarar una función llamada `mergeSort` con un solo parámetro `array`.

```js
assert.isFunction(mergeSort);

const explorer = await __helpers.Explorer(code);
const mergeSortFn = explorer.allFunctions.mergeSort;
assert.exists(mergeSortFn, 'You should declare a mergeSort function');

const params = mergeSortFn.parameters;
assert.lengthOf(
  params,
  1,
  'mergeSort should take a single parameter'
);
assert.isTrue(
  params[0].matches('array'),
  'The parameter of mergeSort should be named array'
);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
