---
id: 69b279bae0f20eda63163cf4
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

Cuando la condición `if` se evalúa como `true`, significa que el elemento en la lista `leftPart` es más pequeño que el elemento con el que se compara en la lista `rightPart`.

En ese caso, puedes colocar el elemento encontrado en `leftArrayIndex` dentro de `leftPart` en el arreglo ordenado.

Dentro del bloque `if`, asigna `leftPart[leftArrayIndex]` a `array[sortedIndex]`.

# --hints--

Debes asignar `leftPart[leftArrayIndex]` a `array[sortedIndex]` dentro de `if`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /array\s*\[\s*sortedIndex\s*\]\s*=\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
