---
id: 69b279bae0f20eda63163cf5
title: Paso 11
challengeType: 1
dashedName: step-11
---

# --description--

Después de asignar el elemento en `leftPart` al arreglo ordenado, incrementa `leftArrayIndex` por `1`.

# --hints--

Debes sumar al valor actual de `leftArrayIndex` después de insertar el elemento en el arreglo ordenado.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /leftArrayIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
