---
id: 69b279bae0f20eda63163cf1
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Ahora es momento de ordenar y fusionar las listas (`leftPart` y `rightPart`) en la lista original `array`.

Puedes hacer esto comparando los elementos de ambas listas y fusionando el elemento más pequeño en la lista principal. Vas a hacer esta comparación para todos los índices en `leftPart` y `rightPart`.

Crea tres variables: `leftArrayIndex`, `rightArrayIndex` y `sortedIndex`, y asígnales el valor `0`. Estas variables te ayudarán a llevar el control de cada índice durante el proceso de ordenamiento.

# --hints--

Debes declarar una variable `leftArrayIndex` inicializada en `0`.

```js
const explorer = await __helpers.Explorer(code);
const { leftArrayIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  leftArrayIndex,
  'You should declare a leftArrayIndex variable inside mergeSort'
);
assert.isTrue(
  leftArrayIndex.value.matches('0'),
  'leftArrayIndex should be initialized to 0'
);
```

Debes declarar una variable `rightArrayIndex` inicializada en `0`.

```js
const explorer = await __helpers.Explorer(code);
const { rightArrayIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  rightArrayIndex,
  'You should declare a rightArrayIndex variable inside mergeSort'
);
assert.isTrue(
  rightArrayIndex.value.matches('0'),
  'rightArrayIndex should be initialized to 0'
);
```

Debes declarar una variable `sortedIndex` inicializada en `0`.

```js
const explorer = await __helpers.Explorer(code);
const { sortedIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  sortedIndex,
  'You should declare a sortedIndex variable inside mergeSort'
);
assert.isTrue(
  sortedIndex.value.matches('0'),
  'sortedIndex should be initialized to 0'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
