---
id: 69b279bae0f20eda63163cf9
title: Paso 15
challengeType: 1
dashedName: step-15
---

# --description--

El ciclo `while` que creaste compara un elemento de `leftPart` con otro de `rightPart`, luego añade el elemento más pequeño a la lista principal `array`.

Continuará esta operación hasta que no queden elementos por comparar. Pero `leftPart` puede seguir teniendo elementos mientras `rightPart` no tiene ninguno, y viceversa.

Crea otro ciclo `while` que se ejecute cuando `leftArrayIndex` sea menor que la longitud de `leftPart`. En los siguientes pasos, lo usarás para copiar los elementos restantes en `leftPart` dentro de la lista `array`.

# --hints--

Debes crear un ciclo `while` con la condición `leftArrayIndex < leftPart.length`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
