---
id: 69b279bae0f20eda63163cf1
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Ahora es momento de ordenar y fusionar las listas (`leftPart` y `rightPart`) en la lista original `array`. 

Puedes hacer esto comparando los elementos de ambas listas y fusionando el elemento más pequeño en la lista principal. Vas a hacer esta comparación para todos los índices en `leftPart` y `rightPart`.

Crea tres variables: `leftArrayIndex`, `rightArrayIndex` y `sortedIndex`, y asígnales el valor `0`. Estas variables te ayudarán a llevar el control de cada índice durante el proceso de ordenamiento.

# --hints--

Debes declarar una variable `leftArrayIndex` inicializada en `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftArrayIndex\s*=\s*0\s*;?/
);
```

Debes declarar una variable `rightArrayIndex` inicializada en `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightArrayIndex\s*=\s*0\s*;?/
);
```

Debes declarar una variable `sortedIndex` inicializada en `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+sortedIndex\s*=\s*0\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);

--fcc-editable-region--
  
--fcc-editable-region--
}
```
