---
id: 66a9ca099e3ecb1be10a2696
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Ahora es el momento de agregar dos estudiantes más a la mesa.

Siguiendo el mismo patrón del paso anterior, añade una cuarto estudiante en una nueva fila de la tabla. Usa los siguientes datos para los elementos de tabla:

- Apellido: `Thompson`
- Nombre: `Jane`
- Grado: `77`

Para la fila correspondiente al quinto estudiante, usa los siguientes datos:

- Apellido: `Williams`
- First Name: `Natalie`
- Grade: `83`

# --hints--

You should have a fourth `tr` inside your `tbody` element.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(4)'));
```

You should have a `td` element with the text `Thompson` inside the fourth `tr` element.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:first-child')?.textContent, 'Thompson');
```

You should have a `td` element with the text `Jane` inside the fourth `tr` element.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(2)')?.textContent, 'Jane');
```

You should have a `td` element with the text `77` inside the fourth `tr` element.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(3)')?.textContent, '77');
```

You should have a fifth `tr` inside your `tbody` element.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 5);
```

You should have a `td` element with the text `Williams` inside the fifth `tr` element.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:first-child')?.textContent, 'Williams');
```

You should have a `td` element with the text `Natalie` inside the fifth `tr` element.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(2)')?.textContent, 'Natalie');
```

You should have a `td` element with the text `83` inside the fifth `tr` element.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(3)')?.textContent, '83');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>
--fcc-editable-region--

--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
