---
id: 66a9b333487b9c14998539a5
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro del elemento `head`, anida un elemento `meta` con el atributo `charset` establecido en el valor `"UTF-8"`.

Debajo de ese elemento `meta`, agrega un elemento `title`.

El texto del elemento `title` debe ser `Calculus Final Exams Table`.

Después de completar tu elemento `head`, puedes agregar tu elemento `body`.

# --hints--

Debes tener un elemento `meta`.

```js
assert.isNotNull(document.querySelector("meta"));
```

El elemento `meta` es un elemento vacío, no debe tener una etiqueta final `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Tu etiqueta `meta` debe tener un atributo `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Tu atributo `charset` debe tener un valor de `"UTF-8"`.

```js
assert.match(code, /charset\s*=\s*('|")UTF-8\1/i);
```

Tu elemento `meta` debe estar anidado dentro del elemento `head`.

```js
const meta = document.querySelector('head > meta');
assert.strictEqual(meta?.parentElement?.tagName, 'HEAD');
```

Debes tener una etiqueta `title` de apertura.

```js
assert.match(code, /<title>/i);
```

Debes tener una etiqueta `title` de cierre.

```js
assert.match(code, /<\/title>/i);
```

Tu elemento `title` debe estar anidado dentro del elemento `head`.

```js
assert.match(code, /<head>.*\s*<title>.*<\/title>.*\s*<\/head>/si);
```

Tu elemento `title` debe tener el texto `Calculus Final Exams Table`. Talvez debas revisar tu ortografía.

```js
const titleText = document.querySelector('title')?.innerText.trim();
assert.strictEqual(titleText, "Calculus Final Exams Table");
```

Debe haber un tag de apertura para `body`.

```js
assert.match(code, /<body>/i);
```

Debe haber un tag de cierre para `body`.

```js
assert.match(code, /<\/body>/i);
```

No debes modificar el elemento `head`. Asegúrese de que no eliminaste tu etiqueta de cierre.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

Tu elemento `body` debe estar después del elemento `head`.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
  --fcc-editable-region--
</html>
```
