---
id: 66a9c006e4f58819396605b3
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Ahora que has completado la sección head, es hora de agregar el <dfn>cuerpo de la tabla</dfn>, `tbody`. El cuerpo de la tabla representará todos los nombres de estudiantes y sus calificaciones.

Agrega un elemento de cuerpo de tabla a tu tabla.

# --hints--

Debes tener una etiqueta de apertura `tbody`.

```js
assert.match(code, /<tbody>/i);
```

Debes tener una etiqueta de cierre `tbody`.

```js
assert.match(code, /<\/tbody>/i);
```

La etiqueta de apertura `tbody` debe venir antes de la etiqueta de cierre `table`.

```js
assert.match(code, /<tbody>[\s\S]*<\/tbody>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>
--fcc-editable-region--

--fcc-editable-region--
    </table>
  </body>
</html>
```
