---
id: 66a9b8f14b963916a3baa732
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Para añadir un título a una `table`, puedes usar el elemento <dfn>table caption</dfn>.

Aquí hay un ejemplo usando el elemento `caption`:

```html
<table>
  <caption>Football Scores</caption>
</table>
```

Dentro de tu elemento `table`, inserta un elemento `caption` con el texto `Calculus Final Exam Grades`.

# --hints--

Deberías tener una etiqueta de apertura `caption`.

```js
assert.match(code, /<caption>/i);
```

Deberías tener una etiqueta de cierre `caption`.

```js
assert.match(code, /<\/caption>/i);
```

Tu elemento `caption` debería estar dentro de tu elemento `table`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.parentElement.tagName, 'TABLE');
```

Tu `caption` debe tener el texto `Calculus Final Exam Grades`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.innerText.trim(), 'Calculus Final Exam Grades');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    --fcc-editable-region--
    <table>

    </table>
    --fcc-editable-region--
  </body>
</html>
```
