---
id: 66a9ca099e3ecb1be10a2696
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Ahora es el momento de agregar dos estudiantes más a la mesa.

Siguiendo el mismo patrón del paso anterior, añade una cuarto estudiante en una nueva fila de la tabla. Usa los siguientes datos para los elementos de tabla:

- Apellido: `Thompson`
- Nombre: `Jane`
- Grado: `77`

Para la fila correspondiente al quinto estudiante, usa los siguientes datos:

- Apellido: `Williams`
- Nombre: `Natalie`
- Grado: `83`

# --hints--

Debes tener una cuarta `tr` dentro de tu elemento `tbody`.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(4)'));
```

Debes tener un elemento `td` con el texto `Thompson` dentro del cuarto elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:first-child')?.textContent, 'Thompson');
```

Debes tener un elemento `td` con el texto `Jane` dentro del cuarto elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(2)')?.textContent, 'Jane');
```

Debes tener un elemento `td` con el texto `77` dentro del cuarto elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(3)')?.textContent, '77');
```

Debes tener una quinta `tr` dentro de tu elemento `tbody`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 5);
```

Debes tener un elemento `td` con el texto `Williams` dentro de la quinta `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:first-child')?.textContent, 'Williams');
```

Debes tener un elemento `td` con el texto `Natalie` dentro de la quinta `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(2)')?.textContent, 'Natalie');
```

Debes tener un elemento `td` con el texto `83` dentro de la quinta `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(3)')?.textContent, '83');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>
--fcc-editable-region--

--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
