---
id: 69aaafeaa4037197c75587cd
title: Eine Queue implementieren
challengeType: 26
dashedName: implement-a-queue
---

# --description--

In diesem Labor implementieren Sie eine Queue-Datenstruktur mithilfe von Funktionen. Eine Queue ist eine First-In-First-Out-(FIFO-)Datenstruktur, bei der Elemente am Ende der Queue hinzugefügt und vorne entfernt werden.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine `enqueue`-Funktion haben, die ein Element am Ende der Queue hinzufügt.
1. Sie sollten eine `dequeue`-Funktion haben, die das vorderste Element der Queue entfernt und zurückgibt.
1. Sie sollten eine `front`-Funktion haben, die das vorderste Element der Queue zurückgibt, ohne es zu entfernen.
1. Sie sollten eine `size`-Funktion haben, die die Anzahl der Elemente in der Queue zurückgibt.
1. Sie sollten eine `isEmpty`-Funktion haben, die `true` zurückgibt, wenn die Queue leer ist, und sonst `false`.

**Hinweis:** Die meisten Tests hängen von der `enqueue`-Funktion ab. Implementieren Sie diese zuerst, da Tests für `dequeue`, `front`, `size` und `isEmpty` das Hinzufügen von Elementen zur Queue erfordern.

# --hints--

Sie sollten eine `enqueue`-Funktion haben.

```js
assert.isFunction(enqueue);
```

Die `enqueue`-Funktion sollte ein Element am Ende der Queue hinzufügen.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

Sie sollten eine `dequeue`-Funktion haben.

```js
assert.isFunction(dequeue);
```

Die `dequeue`-Funktion sollte das vorderste Element der Queue entfernen und zurückgeben.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(dequeue(queue), 'first');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(size(queue), 1);
```

`dequeue` sollte `undefined` zurückgeben, wenn sie auf eine leere Queue angewendet wird.

```js
const queue = initQueue();
assert.isUndefined(dequeue(queue));
```

`enqueue` sollte nach dem Entfernen von Elementen weiterhin Elemente am Ende der Queue hinzufügen.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(dequeue(queue), 'first');
enqueue(queue, 'third');
assert.strictEqual(front(queue), 'second');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(dequeue(queue), 'third');
```

Verschiedene Queue-Instanzen sollten keinen gemeinsamen Zustand teilen.

```js
const queueOne = initQueue();
const queueTwo = initQueue();
enqueue(queueOne, 'first');
enqueue(queueTwo, 'second');
assert.strictEqual(front(queueOne), 'first');
assert.strictEqual(front(queueTwo), 'second');
assert.strictEqual(size(queueOne), 1);
assert.strictEqual(size(queueTwo), 1);
```

Sie sollten eine `front`-Funktion haben.

```js
assert.isFunction(front);
```

Die `front`-Funktion sollte das vorderste Element der Queue zurückgeben, ohne es zu entfernen.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

`front` sollte `undefined` zurückgeben, wenn sie auf eine leere Queue angewendet wird.

```js
const queue = initQueue();
assert.isUndefined(front(queue));
```

Sie sollten eine `size`-Funktion haben.

```js
assert.isFunction(size);
```

Die `size`-Funktion sollte die korrekte Anzahl der Elemente in der Queue zurückgeben.

```js
const queue = initQueue();
assert.strictEqual(size(queue), 0);
enqueue(queue, 'first');
assert.strictEqual(size(queue), 1);
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(size(queue), 3);
dequeue(queue);
assert.strictEqual(size(queue), 2);
```

Sie sollten eine `isEmpty`-Funktion haben.

```js
assert.isFunction(isEmpty);
```

Die `isEmpty`-Funktion sollte `true` für eine leere Queue zurückgeben.

```js
const queue = initQueue();
assert.isTrue(isEmpty(queue));
```

Die `isEmpty`-Funktion sollte `false` für eine nicht-leere Queue zurückgeben.

```js
const queue = initQueue();
enqueue(queue, 'element');
assert.isFalse(isEmpty(queue));
```

Die `isEmpty`-Funktion sollte `true` zurückgeben, nachdem alle Elemente aus der Queue entfernt wurden.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
dequeue(queue);
dequeue(queue);
assert.isTrue(isEmpty(queue));
```

# --seed--

## --seed-contents--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  
}

function dequeue(queue) {
  
}

function front(queue) {
  
}

function size(queue) {
  
}

function isEmpty(queue) {
  
}
```

# --solutions--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  queue.collection.push(element);
}

function dequeue(queue) {
  return queue.collection.shift();
}

function front(queue) {
  return queue.collection[0];
}

function size(queue) {
  return queue.collection.length;
}

function isEmpty(queue) {
  return queue.collection.length === 0;
}
```
