---
id: 66a83bdcf425e7446900b7c4
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

In previous lessons, you learned how to work with the `form` element like this:

```html
<form method="value-goes-here" action="url-goes-here">
  <!-- inputs go inside here -->
</form>
```

`action` 屬性用於指定提交表單時表單數據的發送位置。

`method` 屬性用於指定發送表單數據時使用的 HTTP 方法。 最常見的方法是 `GET` 和 `POST`。

**注意**：你將在後面的模塊中瞭解 HTTP 方法的工作原理。

在你的 `main` 元素內，添加一個 `form` 元素，將 `action` 屬性設置爲 `"https://hotel-feedback.freecodecamp.org"`，`method` 屬性設置爲 `"POST"`。

# --hints--

你應有一個 `form` 開始標籤。

```js
assert.match(code, /<form.*>/i);
```

你應有一個 `form` 結束標籤。

```js
assert.match(code, /<\/form>/i);
```

你的 `form` 元素應有一個 `action` 屬性，設置爲 `"https://hotel-feedback.freecodecamp.org"`。

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('action'), 'https://hotel-feedback.freecodecamp.org');
```

你的 `form` 元素應有一個 `method` 屬性，設置爲 `"POST"`。

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('method'), 'POST');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
    --fcc-editable-region--

    --fcc-editable-region--
    </main>
  </body>
</html>
```
