---
id: 66a8290a27c2c625e2355042
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本講座中，你將通過製作酒店反饋表來練習使用 HTML 表單。

首先添加 `<!DOCTYPE html>`，然後添加具有值爲 `en` 的 `lang` 屬性的 `html` 元素。

在 `html` 元素中，添加一個 `head` 元素。


# --hints--

你應該有 `<!DOCTYPE html>`。

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

你應該有一個 `html` 開始標籤，語言設置爲英語。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

你應該有一個 `html` 結束標籤。

```js
assert.match(code, /<\/html>/i);
```

你的 `DOCTYPE` 應放在 `html` 元素之前。

```js
assert.match(code, /<!DOCTYPE\s+html>[.\n\s]*<html\s+lang\s*=\s*('|")en\1\s*>/im)
```

你應該有一個 `head` 開始標籤。

```js
assert.match(code, /<head>/i);
```

你應該有一個 `head` 結束標籤。

```js
assert.match(code, /<\/head>/i);
```

你的 `head` 開始標籤應在 `head` 結束標籤之前。

```js
assert.match(code, /<head>[.\n\s]*<\/head>/im)
```

你的 `head` 元素應該在 `html` 元素中。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[.\n\s]*<head>[.\n\s]*<\/head>[.\n\s]*<\/html>/im)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
