---
id: 66a83e5e491625454b6f62c3
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

表單由 `inputs` 組成，使用者可以在其中輸入他們的資料。你可以使用 `fieldset` 元素將相關的 inputs 群組在一起。 

以下是使用 `fieldset` 元素的範例：

```html
<form action="/example-url">
  <fieldset>
  <!-- inputs go inside here-->
  </fieldset>
</form>
```

在你的 `form` 元素內，添加一個 `fieldset` 元素。

# --hints--

你應該有一個開啟的 `fieldset` 頁籤。

```js
assert.match(code, /<fieldset>/i);
```

你應該有一個結束的 `fieldset` 頁籤。

```js
assert.match(code, /<\/fieldset>/i);
```

你的 `fieldset` 元素應該在 `form` 元素內。

```js
assert.isNotNull(document.querySelector('form fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">        
--fcc-editable-region--
        
--fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
