---
id: 66a83fec026a7a4631e084d2
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

在使用 `fieldset` 元素時，通常會使用標題來描述輸入的群組。你可以使用 `legend` 元素來達成這個目的。

以下是使用 `legend` 元素的範例：

```html
<form action="/example-url">
  <fieldset>
    <legend>Personal Information</legend>
    <!-- inputs go inside here-->
  </fieldset>
</form>
```

在你的 `fieldset` 元素內，為你的 `legend` 元素添加文字 `Personal Information`。

# --hints--

你應該有一個開啟的 `legend` 頁籤。

```js
assert.match(code, /<legend>/);
```

你應該有一個結束的 `legend` 頁籤。

```js
assert.match(code, /<\/legend>/);
```

你的 `legend` 元素應該在 `fieldset` 元素內。

```js
assert.isNotNull(document.querySelector('fieldset legend'));
```

你的 `legend` 元素應包含文字 `Personal Information`。

```js
const legend = document.querySelector('fieldset legend');
assert.strictEqual(legend?.innerText.trim(), 'Personal Information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
