---
id: 66a83601cd819e37f0dccd14
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

在這個工作坊中，你將透過建置一個 Hotel Feedback Form 來練習操作 HTML 表單。

對於介紹文字，你會想要顯示主標題，接著是一則關於留下回授的簡短說明。 

在你的 `body` 元素內，為你的 `body` 元素添加一個 `header` 元素。

在 `header` 元素內，添加一個 `h1` 元素。`h1` 元素的文字應該是 `Hotel Feedback Form`。

在你的 `h1` 元素下方，添加一個 `p` 元素。`p` 元素的文字應該是 `Thank you for staying with us. Please provide feedback on your recent stay.`

# --hints--

你應該有一個開啟的 `header` 頁籤。

```js
assert.match(code, /<header>/i);
```

你應該有一個結束的 `header` 頁籤。

```js
assert.match(code, /<\/header>/i);
```

你應該有一個開啟的 `h1` 頁籤。

```js
assert.match(code, /<h1>/i);
```

你應該有一個結束的 `h1` 頁籤。

```js
assert.match(code, /<\/h1>/i);
```

你的 `h1` 元素應該有文字 `Hotel Feedback Form`。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.innerText.trim(), 'Hotel Feedback Form');
```

你的 `h1` 元素應該在 `header` 元素內。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.parentElement.tagName, 'HEADER');
```

你應該有一個開啟的 `p` 頁籤。

```js
assert.match(code, /<p>/i);
```

你應該有一個結束的 `p` 頁籤。

```js
assert.match(code, /<\/p>/i);
```

你的段落元素應該包含文字 `Thank you for staying with us. Please provide feedback on your recent stay.`

```js
const pElement = document.querySelector('header p');
assert.strictEqual(pElement?.innerText.trim(), 'Thank you for staying with us. Please provide feedback on your recent stay.');
```

你的段落元素應該在你的 `header` 裡面。

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.parentElement.tagName, 'HEADER');
```

你的段落元素應該放在你的 `h1` 之後。

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.previousElementSibling.tagName, 'H1');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
