---
id: 66a83601cd819e37f0dccd14
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

對於介紹性文字，你需要顯示主標題，然後是關於留下反饋意見的簡短說明。

在 `body` 元素內，添加一個 `header` 元素。

在 `header` 元素中，添加一個 `h1` 元素。 `h1` 元素的文本應爲 `Hotel Feedback Form`。

在你的 `h1` 元素下方，添加一個 `p` 元素。 `p` 元素的文本應爲 `Thank you for staying with us. Please provide feedback on your recent stay.`

# --hints--

你應該有一個 `header` 開始標籤。

```js
assert.match(code, /<header>/i);
```

你應該有一個 `header` 結束標籤。

```js
assert.match(code, /<\/header>/i);
```

你應該有一個 `h1` 開始標籤。

```js
assert.match(code, /<h1>/i);
```

你應該有一個 `h1` 結束標籤。

```js
assert.match(code, /<\/h1>/i);
```

你的 `h1` 元素應包含文本 `Hotel Feedback Form`。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.innerText.trim(), 'Hotel Feedback Form');
```

你的 `h1` 元素應該在 `header` 元素中。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.parentElement.tagName, 'HEADER');
```

你應該有一個 `p` 開始標籤。

```js
assert.match(code, /<p>/i);
```

你應該有一個 `p` 結束標籤。

```js
assert.match(code, /<\/p>/i);
```

你的段落元素應包含以下文本 `Thank you for staying with us. Please provide feedback on your recent stay.`

```js
const pElement = document.querySelector('header p');
assert.strictEqual(pElement?.innerText.trim(), 'Thank you for staying with us. Please provide feedback on your recent stay.');
```

你的段落元素應位於 `header` 內。

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.parentElement.tagName, 'HEADER');
```

你的段落元素應位於 `h1` 之後。

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.previousElementSibling.tagName, 'H1');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
