---
id: 66a83bdcf425e7446900b7c4
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

在之前的課程中，你學會了如何這樣操作 `form` 元素：

```html
<form method="value-goes-here" action="url-goes-here">
  <!-- inputs go inside here -->
</form>
```

`action` 屬性用於指定當表單提交時，表單資料應該送往的位置。

`method` 屬性用於指定在傳送表單資料時要使用的 HTTP 方法。最常見的方法是 `GET` 和 `POST`。 

**注意**：你將在後續模組中學習 HTTP 方法如何運作。

在你的 `main` 元素內，為你的 `form` 元素添加一個 `action` 屬性，設定為 `"https://hotel-feedback.freecodecamp.org"`，並添加一個 `method` 屬性，設定為 `"POST"`。

# --hints--

你應該有一個開啟的 `form` 頁籤。

```js
assert.match(code, /<form.*>/i);
```

你應該有一個結束的 `form` 頁籤。

```js
assert.match(code, /<\/form>/i);
```

你的 `form` 元素應該有一個 `action` 屬性設定為 `"https://hotel-feedback.freecodecamp.org"`。

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('action'), 'https://hotel-feedback.freecodecamp.org');
```

你的 `form` 元素應該有一個 `method` 屬性設定為 `"POST"`。

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('method'), 'POST');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
