---
id: 66a84111965a0c46df6bbd0a
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

在之前的課程中，你學會了如何將 `label` 元素與 `input` 這樣關聯：

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

`label` 元素上的 `for` 屬性應該與 `input` 元素上的 `id` 屬性相符。這稱為明確的關聯。

在你的 `legend` 元素下方，為你的 `label` 元素添加文字 `Name (required):`。將它的 `for` 屬性設為 `"full-name"`。

然後在你的 `label` 元素下方，添加一個沒有屬性的 `input` 元素。在接下來的步驟中，你將添加必要的屬性。 

# --hints--

你應該有一個開啟的 `label` 頁籤。

```js
assert.match(code, /<label.*>/);
```

你應該有一個結束的 `label` 頁籤。

```js
assert.match(code, /<\/label>/);
```

你的 `label` 元素應該放置在 `fieldset` 元素內的 `legend` 下面。

```js
assert.isNotNull(document.querySelector('fieldset legend + label'));
```

你的 `label` 元素應該有一個 `for` 屬性，值為 `"full-name"`。

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.getAttribute('for'), 'full-name');
```

你的 `label` 元素應該有 `Name (required):` 的文字。

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.innerText.trim(), 'Name (required):');
```

你應該在你的 `label` 元素下方放置一個 `input` 元素。

```js
assert.isNotNull(document.querySelector('fieldset label + input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
