---
id: 66a83fec026a7a4631e084d2
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

在使用 `fieldset` 元素時，通常使用一個標題來描述輸入組。 爲此，你可以使用 `legend` 元素。

下面是使用 `legend` 元素的示例：

```html
<form action="/example-url">
  <fieldset>
    <legend>Personal Information</legend>
    <!-- inputs go inside here-->
  </fieldset>
</form>
```

在你的 `fieldset` 元素中，添加一個 `legend` 元素，文本爲 `Personal Information`。

# --hints--

你應該有一個 `legend` 開始標籤。

```js
assert.match(code, /<legend>/);
```

你應該有一個 `legend` 結束標籤。

```js
assert.match(code, /<\/legend>/);
```

你的 `legend` 元素應該在 `fieldset` 元素中。

```js
assert.isNotNull(document.querySelector('fieldset legend'));
```

你的 `legend` 元素應包含文本 `Personal Information`。

```js
const legend = document.querySelector('fieldset legend');
assert.strictEqual(legend?.innerText.trim(), 'Personal Information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
--fcc-editable-region--
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>

        </fieldset>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>
```
