---
id: 66a84111965a0c46df6bbd0a
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

In previous lessons, you learned how to associate a `label` element with an `input` like this:

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

`label` 元素的 `for` 屬性應與對應的 `input` 元素的 `id` 屬性匹配。 這稱爲顯式關聯。

在你的 `legend` 元素下方，添加一個 `label`元素，其文本爲 `Name (required):`。 將其 `for` 屬性設置爲 `"full-name"`。

然後在你的 `label` 元素下方，添加一個無屬性的 `input` 元素。 在接下來的步驟中，你將添加必要的屬性。

# --hints--

你應該有一個 `label` 開始標籤。

```js
assert.match(code, /<label.*>/);
```

你應該有一個 `label` 結束標籤。

```js
assert.match(code, /<\/label>/);
```

你的 `label` 元素應放在 `fieldset` 元素中 `legend` 的下方。

```js
assert.isNotNull(document.querySelector('fieldset legend + label'));
```

你的 `label` 元素應該有一個 `for` 屬性，值爲 `"full-name"`。

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.getAttribute('for'), 'full-name');
```

你的 `label` 元素的文本應爲 `Name (required):`。

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.innerText.trim(), 'Name (required):');
```

你應該在 `label` 元素下方有一個 `input` 元素。

```js
assert.isNotNull(document.querySelector('fieldset label + input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
--fcc-editable-region--

--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
