---
id: 66a8380d911e3f4270d5cadc
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

現在，是時候添加表述頁面主要內容的 `main` 元素。

# --hints--

你應該有一個開啟的 `main` 頁籤。

```js
assert.match(code, /<main>/i);
```

你應該有一個結束的 `main` 頁籤。

```js
assert.match(code, /<\/main>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
