---
id: 66a9401c9d660d6bb15993e2
title: 步驟 16
challengeType: 0
dashedName: step-16
---

# --description--

你的酒店反饋表應爲用戶添加年齡選項。

首先，在表單中添加一個 `label` 元素，文本爲 `Age (optional):`。

`for` 屬性應設置爲 `"age"`。

# --hints--

你應該有一個 `label` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 3);
```

你的 `label` 元素應將 `for` 屬性設置爲 `"age"`。

```js
assert.isNotNull(document.querySelector('label[for="age"]'));
```

你的 `label` 元素應包含文本 `Age (optional):`。

```js
assert.strictEqual(document.querySelector('label[for="age"]')?.textContent.trim(), "Age (optional):");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
--fcc-editable-region--

--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
