---
id: 66a930b20f589b6664c51cb0
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

當使用者提供他們的全名時，`input` 將接受純文字。

在之前的課程中，你學會了如何這樣使用 `type` 屬性：

```html
<input type="text">
```

為你現有的 `input` 元素添加一個 `type` 屬性，並設定為 `"text"`。 

同時為 `input` 元素添加一個值為 `"full-name"` 的 `id` 屬性。

# --hints--

你不應該移除現有的 `input` 元素。

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

你不應該創建新的 `input` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

你的 `input` 元素應該有一個 `type` 屬性，其值為 `"text"`。

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('type'), 'text');
```

你的 `input` 元素應該有一個 `id` 屬性，值為 `"full-name"`。

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('id'), 'full-name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
        --fcc-editable-region--
          <input>
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
