---
id: 66a954b2bcddba72076c1857
title: Step 17
challengeType: 0
dashedName: step-17
---

# --description--

在你的 `label` 元素下方，添加一個 `radio` 按鈕，`id` 設置爲 `"no-option"`，名字屬性設置爲 `"hotel-stay"`。

在你的新 `radio` 按鈕下方，添加另一個 `label` 元素，其 `for` 屬性設置爲 `"no-option"`。 `label` 的文本應爲 `No`。

完成後，你就可以嘗試使用單選按鈕，每次選擇一個選項。

# --hints--

你應該在第一個 `label` 元素下方有一個 `input` 元素。

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) label + input'));
```

你的 `input` 元素應該有一個 `type` 屬性，值爲 `radio`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.getAttribute('type'), 'radio');
```

你的 `input` 元素應該有一個 `id` 屬性，值爲 `"no-option"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.id, 'no-option');
```

你的 `input` 元素應該有一個 `name` 屬性，值爲 `"hotel-stay"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.name, 'hotel-stay');
```

你的第二個 `input` 元素下方應該有第二個 `label` 元素。

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label'));
```

你的 `label` 元素應該有一個 `for` 屬性，設置爲 `"no-option"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.htmlFor, "no-option");
```

你的 `label` 元素應該包含文本 `No`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.textContent?.trim(), 'No');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          --fcc-editable-region--

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
