---
id: 66a9521bc70162712caf118d
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--

如果想讓用戶從一個選項列表中選擇一個選項，可以使用一組單選按鈕。

下面是兩個單選按鈕的示例：

```html
<input type="radio" id="yes" name="first-time">
<label for="yes">Yes</label>
<input type="radio" id="no" name="first-time">
<label for="no">No</label>
```

在此示例中，單選按鈕通過使用相同的 `name` 屬性值分組在一起。 這意味着一次只能選擇一個單選按鈕。

在你的 `legend` 元素下方，添加一個 `radio` 按鈕，`id` 設置爲 `"yes-option"`，`name` 屬性設置爲 `"hotel-stay"`。

在你的 `radio` 按鈕下方，添加一個 `label` 元素，文本爲 `Yes`，`for` 屬性設置爲 `"yes-option"`。

# --hints--

你的 `legend` 元素下方應該有一個 `radio` 按鈕。

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]'));
```

你的 `radio` 按鈕應將 `id` 屬性設置爲 `"yes-option"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.id, 'yes-option');
```

你的 `radio` 按鈕應將 `name` 屬性設置爲 `"hotel-stay"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.name, 'hotel-stay');
```

你應該在你的 `radio` 按鈕下方有一個 `label` 元素。

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label'));
```

你的 `label` 元素應將 `for` 屬性設置爲 `"yes-option"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.htmlFor, 'yes-option');
```

你的 `label` 元素應該包含文本 `Yes`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.innerText?.trim(), 'Yes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        --fcc-editable-region--
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>

        </fieldset>
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
