---
id: 66a954b2bcddba72076c1857
title: 步驟 20
challengeType: 0
dashedName: 步驟-20
---

# --description--

在第一個單選按鈕下方，添加另一個 `label` 元素，其 `for` 屬性設置爲 `"no-option"`。 `label` 的文本應爲 `No`。

在第二個 `label` 元素下方，添加一個 `radio` 按鈕，`id` 設置爲 `"no-option"`，`name` 屬性設置爲 `"hotel-stay"`。

完成後，你就可以嘗試使用單選按鈕，每次選擇一個選項。

# --hints--

你應該在第一個單選按鈕下方添加一個 `label` 元素。

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label'));
```

你的 `label` 應將 `for` 屬性設置爲 `"no-option"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.htmlFor, "no-option");
```

你的 `label` 應該包含文本 `No`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label[for="no-option"]')?.textContent, 'No');
```

你的 `label` 下方應該有第二個 `input`。

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(2) label + input'), 2);
```

你的 `input` 應該有一個 `type`，值爲 `radio`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input:nth-of-type(2)')?.getAttribute('type'), 'radio');
```

你的 `input` 應該有一個 `id`，值爲 `"no-option"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input:nth-of-type(2)')?.id, 'no-option');
```

你的 `input` 應該有一個 `name`，值爲 `"hotel-stay"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input:nth-of-type(2)')?.name, 'hotel-stay');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          --fcc-editable-region--
          <label for="yes-option">Yes</label>
          <input id="yes-option" type="radio" name="hotel-stay" />

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
