---
id: 66a93c95bc58e26a8fe95818
title: 步驟 14
challengeType: 0
dashedName: step-14
---

# --description--

接下來，在你的郵箱 `label` 下方，添加一個 `input`，類型爲 `"email"`。 該 `input` 元素應該有一個 `id` 屬性，值設置爲 `"email"`。 將 `name` 屬性也設置爲 `"email"`。

該 `input`也是必填的，因此，請確保添加 `required` 屬性。

最後，添加一個 `placeholder` 屬性，設置爲 `"example@email.com"`。

# --hints--

你的郵箱 `label` 元素下面應該有一個 `input` 元素。

```js
assert.isNotNull(document.querySelector('label[for="email"] + input'));
```

你的 `input` 元素應該有一個 `type`，值爲 `"email"`。

```js
assert.isNotNull(document.querySelector('input[type="email"]'));
```

你的 `input` 元素應有一個 `id` 屬性，設置爲 `"email"`。

```js
assert.isNotNull(document.querySelector('input#email'));
```

你的 `input` 元素應有一個 `name` 屬性，設置爲 `"email"`。

```js
assert.isNotNull(document.querySelector('input[name="email"]'));
```

你的 `input` 元素應該具有 `required` 屬性。

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('required'), "");
```

您的 `輸入的` 元素應該有一個 `佔位符` 屬性設置爲 `"example@email.com"`

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('placeholder'), "example@email.com");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>
        --fcc-editable-region--
          <label for="email">Email address (required):</label>

        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
