---
id: 66a93730719e1f68410cce54
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

`name` 屬性用於在表單數據提交到服務器後對其進行標識。

下面是一個如何使用 `name` 屬性的示例：

```html
<input type="email" name="email">
```

爲 `input` 元素添加一個 `name` 屬性，值爲 `"name"`。

# --hints--

你不應刪除現有的 `input` 元素。

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

你不應創建新的 `input` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

你的 `input` 元素應該有一個 `name` 屬性，值爲 `"name"`。

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('name'), 'name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
