---
id: 66a93bbe65a26169dbf3bc39
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

你的飯店回授表單也應該從使用者收集電子郵件位址。 

首先為表單添加一個新的 `label` 元素，並將文字設為 `Email address (required):`。你的 `label` 元素應該有一個 `for` 屬性，設定為 `"email"`。

# --hints--

你應該有一個 `label` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 2);
```

你的 `label` 元素應該有文字 `Email address (required):`。

```js
assert.strictEqual(document.querySelector('fieldset label:nth-of-type(2)')?.innerText.trim(), 'Email address (required):');
```

你的 `label` 元素應該有一個 `for` 屬性，設定為 `"email"` 的值。

```js
assert.strictEqual(document.querySelector('fieldset label[for="email"]')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
