---
id: 66a94ea5df66236ebed260e8
title: 第 15 步
challengeType: 0
dashedName: step-15
---

# --description--

表單中的下一個區段將負責詢問使用者是否曾經入住過該飯店。 

首先為你的文件添加一個 `fieldset` 元素。

在 `fieldset` 元素內，為你的 `legend` 元素添加文字 `Was this your first time at our hotel?`。 

# --hints--

你應該有一個 `fieldset` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 2);
```

你的 `fieldset` 元素內應該有一個 `legend` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset legend'), 2);
```

你的 `legend` 元素應該包含 `Was this your first time at our hotel?` 的文字。

```js
assert.strictEqual(document.querySelectorAll('fieldset legend')[1]?.textContent.trim(), 'Was this your first time at our hotel?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
