---
id: 66a962954f4e0d76223b37ed
title: 步驟 24
challengeType: 0
dashedName: step-24
---

# --description--

接下來，添加另一個複選框 `input`，其 `id`、`name` 和 `value` 屬性都設置爲 `"location"`。

此 `label` 元素的文本爲 `Location`， `for` 屬性應設置爲 `"location"`。

在該 `label` 元素下方，添加另一個複選框 `input`，其 `id`、`name` 和 `value` 屬性都設置爲 `"reputation"`。

此 `label` 元素的文本爲 `Reputation`， `for` 屬性應設置爲 `"reputation"`。

# --hints--

你應該有第三個複選框 `input`。

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)"));
```

你的第三個複選框 `input` 應該有一個 `id` 屬性，設置爲 `"location"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('id'), 'location');
```

你的第三個複選框 `input` 應該有一個 `name` 屬性，設置爲 `"location"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('name'), 'location');
```

你的第三個複選框 `input` 應該有一個 `value` 屬性，設置爲 `"location"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('value'), 'location');
```

你應該在你的複選框 `input` 之後有一個 `label` 元素，文本爲 `Location`。

```js 
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[value='location'] + label")?.textContent?.trim(), 'Location');
```

你的 `label` 應將 `for` 屬性設置爲 `"location"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="location"] + label')?.getAttribute('for'), 'location');
```

你的 `label` 下方應該有第四個複選框 `input`。

```js 
assert.exists(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']"));
```

你的第四個複選框 `input` 應該有一個 `id` 屬性，設置爲 `"reputation"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('id'), 'reputation');
```

你的第四個複選框 `input` 應該有一個 `name` 屬性，設置爲 `"reputation"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('name'), 'reputation');
```

你的第四個複選框 `input` 應該有一個 `value` 屬性，設置爲 `"reputation"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('value'), 'reputation');
```

在聲譽複選框下方，你應該有一個 `label` 元素，文本爲 `Reputation`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.textContent?.trim(), 'Reputation');
```

你的 `label` 應將 `for` 屬性設置爲 `"reputation"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.getAttribute('for'), 'reputation');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="ads" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="recommendation"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          --fcc-editable-region--

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
