---
id: 66a930b20f589b6664c51cb0
title: 步驟 10
challengeType: 0
dashedName: step-10
---

# --description--

當用戶提供全名時，`input` 將接受純文本。

In previous lessons, you learned how to work with the `type` attribute like this:

```html
<input type="text">
```

對於你現有的 `input` 元素，添加一個 `type` 屬性，設置爲 `"text"`。

也爲 `input` 元素賦予一個 `id` 屬性，其值爲 `"full-name"`。

# --hints--

你不應刪除現有的 `input` 元素。

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

你不應創建新的 `input` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

你的 `input` 元素應該有一個 `type` 屬性，值爲 `"text"`。

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('type'), 'text');
```

你的 `input` 元素應該有一個 `id` 屬性，值爲 `"full-name"`。

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('id'), 'full-name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
        --fcc-editable-region--
          <input>
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
