---
id: 66a9419e2d18476c645ce693
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

數字 `input` 用於創建數字輸入框。

下面是一個數字輸入框的示例：

```html
<input type="number" id="age" name="age" min="18" max="100">
```

`min` 和 `max` 屬性用於設置輸入字段中可輸入的最小值和最大值。

在你的 `label` 元素下方，添加一個輸入，`type` 屬性設置爲 `"number"`，`id` 爲 `"age"`。

`name` 屬性應設置爲 `"age"`，`min` 屬性應設置爲 `"3"`，`max` 屬性應設置爲 `"100"`。

# --hints--

你的 `label` 元素下方應該有一個數字 `input`。

```js
assert.isNotNull(document.querySelector('label + input[type="number"]'));
```

你的數字 `input` 應該有一個 `id`，值爲 `age`。

```js
assert.isNotNull(document.querySelector('label + input#age'));
```

你的數字 `input` 應有一個 `name` 屬性，設置爲 `age`。

```js
assert.isNotNull(document.querySelector('label + input[name="age"]'));
```

你的數字 `input` 應有一個 `min` 屬性，設置爲 `"3"`。

```js
assert.isNotNull(document.querySelector('label + input[min="3"]'));
```

你的數字 `input` 應有一個 `max` 屬性，設置爲 `"100"`。

```js
assert.isNotNull(document.querySelector('label + input[max="100"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
      --fcc-editable-region--
          <label for="age">Age (optional):</label>

      --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
