---
id: 66a937e74920ba68ebe5e86d
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

In previous lessons, you learned how to work with the `placeholder` and `required` attributes like this:

```html
<input type="text" placeholder="e.g., John Doe" required>
```

For your existing `input` element, add a `placeholder` attribute with the value of `"e.g., John Doe"`.

也在 `input` 元素中添加 `required` 屬性。

# --hints--

你不應刪除現有的 `input` 元素。

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Your `input` should have a `placeholder` attribute with the value of `"e.g., John Doe"`.

```js
assert.strictEqual(document.querySelector('fieldset input')?.getAttribute('placeholder'), 'e.g., John Doe');
```

你的 `input` 應該具有 `required` 屬性。

```js
assert.isNotNull(document.querySelector('fieldset input[required]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
