---
id: 66a972137acd1179fa3fe8a0
title: 步驟 26
challengeType: 0
dashedName: step-26
---

# --description--

在你的 `select` 元素內，為你的 `option` 元素添加以下五個 `option` 元素，並為 `option` 文字和 `value` 屬性設定相對應的值：

**值屬性：**

- 差
- 令人滿意的
- 良好
- 非常好
- 優秀

**選項元素文字：**

- 差
- 令人滿意
- 良好
- 非常好
- 優秀

# --hints--

你應該有一個 `option` 元素，其值設定為 `"poor"`。

```js
assert.exists(document.querySelector('option[value="poor"]'));
```

你的 `option` 元素中 `value` 為 `"poor"` 的應該有 `Poor` 這段文字。

```js
assert.strictEqual(document.querySelector('option[value="poor"]')?.textContent.trim(), 'Poor');
```

你應該有一個 `option` 元素，其 `value` 設定為 `"satisfactory"`。

```js
assert.exists(document.querySelector('option[value="satisfactory"]'));
```

你的 `option` 元素中 `value` 為 `"satisfactory"` 的應該有 `Satisfactory` 這段文字。

```js
assert.strictEqual(document.querySelector('option[value="satisfactory"]')?.textContent.trim(), 'Satisfactory');
```

你應該有一個 `option` 元素，其 `value` 設定為 `"good"`。

```js
assert.exists(document.querySelector('option[value="good"]'));
```

你的 `option` 元素中 `value` 為 `"good"` 的應該有 `Good` 這段文字。

```js

assert.strictEqual(document.querySelector('option[value="good"]')?.textContent.trim(), 'Good');
```

你應該有一個 `option` 元素，其值設為 `"very-good"`。

```js
assert.exists(document.querySelector('option[value="very-good"]'));
```

你的 `option` 元素中 `value` 為 `"very-good"` 的應該有 `Very Good` 的文字。

```js
assert.strictEqual(document.querySelector('option[value="very-good"]')?.textContent.trim(), 'Very Good');
```

你應該有一個 `option` 元素，其值設定為 `"excellent"`。

```js
assert.exists(document.querySelector('option[value="excellent"]'));
```

你的 `option` 元素中 `value` 為 `"excellent"` 的應該有 `Excellent` 的文字。

```js
assert.strictEqual(document.querySelector('option[value="excellent"]')?.textContent.trim(), 'Excellent');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          <input type="checkbox" id="location" name="choice" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="choice"
            value="reputation"
          />
          <label for="reputation">Reputation</label>
          
          <input type="checkbox" id="price" name="choice" value="price" />
          <label for="price">Price</label>
        </fieldset>

        <fieldset>
          <legend>Ratings</legend>

          <label for="service">How was the service?</label>

          <select name="service" id="service">
          --fcc-editable-region--
            
          --fcc-editable-region--
          </select>
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
