---
id: 66a93c95bc58e26a8fe95818
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

接著，在你的 email `label` 下方添加一個型別為 `"email"` 的 `input`。這個 `input` 應該有一個 `id` 屬性，設定為 `"email"` 的值。同時，將 `name` 屬性設定為 `"email"` 的值。 

這個 `input` 也是必填的，所以請確保為你的 `input` 添加 `required` 屬性。

最後，為你的元素添加一個 `placeholder` 屬性，並設為 `"example@email.com"`。

# --hints--

你應該在你的電子郵件 `標籤` 下方有一個 `input` 元素。

```js
assert.isNotNull(document.querySelector('label[for="email"] + input'));
```

你的 `input` 元素的 `type` 應該是 `"email"`。

```js
assert.isNotNull(document.querySelector('input[type="email"]'));
```

你的 `input` 元素應該有一個 `id` 屬性設定為 `"email"`。

```js
assert.isNotNull(document.querySelector('input#email'));
```

你的 `input` 元素應該有一個 `name` 屬性設定為 `"email"`。

```js
assert.isNotNull(document.querySelector('input[name="email"]'));
```

你的 `input` 元素應該有一個 `required` 屬性。

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('required'), "");
```

你的 `input` 元素應該有一個 `placeholder` 屬性，設定為 `"example@email.com"`。
  
```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('placeholder'), "example@email.com");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
