---
id: 66a9521bc70162712caf118d
title: 第 16 步
challengeType: 0
dashedName: step-16
---

# --description--

如果你想讓使用者從一個選項列表中選擇一個選項，你可以使用一組圓鈕。

這裡有兩個圓鈕的範例：

```html
<input type="radio" id="yes" name="first-time">
<label for="yes">Yes</label>
<input type="radio" id="no" name="first-time">
<label for="no">No</label>
```

在此範例中，圓鈕是透過使用相同的 `name` 屬性值來群組在一起。這表示一次只能選取一個圓鈕。

在你的 `legend` 元素下方，為你的 `radio` 按鈕添加 `id` 設為 `"yes-option"`、`name` 屬性設為 `"hotel-stay"`，以及 `value` 屬性設為 `"yes"`。

在你的 `radio` 按鈕下方，為你的 `label` 元素添加文字 `Yes` 並將 `for` 屬性設為 `"yes-option"`。

# --hints--

你應該在你的 `legend` 元素下方放置一個 `radio` 按鈕。

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]'));
```

你的 `radio` 按鈕應該有一個 `id` 屬性設定為 `"yes-option"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.id, 'yes-option');
```

你的 `radio` 按鈕應該設定 `name` 屬性為 `"hotel-stay"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.name, 'hotel-stay');
```

你的圓鈕應該有一個 `value` 屬性設定為 `"yes"`。

```js
assert.strictEqual(
  document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.value,
  'yes'
);
```

你應該在你的 `radio` 按鈕下方放置一個 `label` 元素。

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label'));
```

你的 `label` 元素應該有一個 `for` 屬性設為 `"yes-option"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.htmlFor, 'yes-option');
```

你的 `label` 元素應該有文字 `Yes`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.innerText?.trim(), 'Yes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
