---
id: 66a96127422411756204bc1b
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

添加另一個複選框 `input`，其 `id`、`name` 和 `value` 屬性都設置爲 `"recommendation"`。

在該複選框 `input` 之後，添加另一個 `label`，文本爲 `Personal Recommendation`。 `for` 屬性應設置爲 `"recommendation"`。

# --hints--

你的 `label` 元素下方應該有一個複選框 `input`。

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(3) label + input[type="checkbox"]'));
```

你的複選框 `input` 應該有一個 `id` 屬性，設置爲 `"recommendation"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']")?.id, "recommendation");
```

你的複選框 `input` 應該有一個 `name` 屬性，設置爲 `"recommendation"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']")?.name, "recommendation");
```

你的複選框 `input` 應該有一個 `value` 屬性，設置爲 `"recommendation"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']")?.value, "recommendation");
```

你的複選框  `input`  下方應該有一個 `label` 元素。

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label:nth-of-type(2)"));
```

你的 `label` 元素的文本應爲 `Personal Recommendation`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[type="checkbox"] + label:nth-of-type(2)')?.textContent?.trim(), "Personal Recommendation");
```

你的 `label` 元素該有一個 `for` 屬性，設置爲 `"recommendation"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[type="checkbox"] + label:nth-of-type(2)')?.getAttribute("for"), "recommendation");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="ads" value="ads" />
          <label for="ads">Social Media Ads</label>
          --fcc-editable-region--

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
