---
id: 66a962954f4e0d76223b37ed
title: 第 21 步
challengeType: 0
dashedName: step-21
---

# --description--

接著，為你的表單添加另一個勾選框 `input`，其 `id` 和 `value` 屬性設為 `"location"`，`name` 屬性設為 `"choice"`。

對於 `label` 元素，`Location` 的文字和 `for` 屬性應該設為 `"location"`。

在該 `label` 元素下方，添加另一個核取方塊 `input`，其 `id` 和 `value` 屬性設為 `"reputation"`，`name` 屬性設為 `"choice"`。

對於 `label` 元素，`Reputation` 的文字和 `for` 屬性應設為 `"reputation"`。

# --hints--

你應該有第三個核取方塊 `input`。

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)"));
```

你的第三個核取方塊 `input` 應該有一個 `id` 屬性設為 `"location"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('id'), 'location');
```

你的第三個核取方塊 `input` 應該有一個 `name` 屬性設為 `"choice"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('name'), 'choice');
```

你的第三個核取方塊 `input` 應該有一個 `value` 屬性設定為 `"location"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('value'), 'location');
```

你應該在你的核取方塊 `input` 後面放置一個 `label` 元素，並帶有文字 `Location`。

```js 
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[value='location'] + label")?.textContent?.trim(), 'Location');
```

你的 `label` 應該將 `for` 屬性設為 `"location"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="location"] + label')?.getAttribute('for'), 'location');
```

你應該在你的 `label` 下方有第四個核取方塊 `input`。

```js 
assert.exists(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']"));
```

你應該有第四個核取方塊 `input`，其 `id` 屬性設定為 `"reputation"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('id'), 'reputation');
```

你的第四個核取方塊 `input` 應該有一個 `name` 屬性設為 `"choice"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('name'), 'choice');
```

你的第四個核取方塊 `input` 應該有一個 `value` 屬性設定為 `"reputation"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('value'), 'reputation');
```

你應該在聲譽核取方塊下方有一個帶有文字 `Reputation` 的 `label` 元素。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.textContent?.trim(), 'Reputation');
```

你的 `label` 應該將 `for` 屬性設為 `"reputation"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.getAttribute('for'), 'reputation');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
