---
id: 66a9401c9d660d6bb15993e2
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

你的飯店回授表單應該有一個選項，可讓使用者添加他們的年齡。 

首先為表單添加一個帶有文字 `Age (optional):` 的 `label` 元素。

`for` 屬性應該設定為 `"age"`。

# --hints--

你應該有一個 `label` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 3);
```

你的 `label` 元素應該有一個 `for` 屬性設為 `"age"`。

```js
assert.isNotNull(document.querySelector('label[for="age"]'));
```

你的 `label` 元素應該有文字 `Age (optional):`。

```js
assert.strictEqual(document.querySelector('label[for="age"]')?.textContent.trim(), "Age (optional):");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
