---
id: 66a937e74920ba68ebe5e86d
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

在之前的課程中，你學會了如何這樣使用 `placeholder` 和 `required` 屬性：

```html
<input type="text" placeholder="e.g., John Doe" required>
```

為你的現有 `input` 元素添加一個 `placeholder` 屬性，值為 `"e.g., John Doe"`。

此外，為你的 `input` 元素添加 `required` 屬性。

# --hints--

你不應該移除現有的 `input` 元素。

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

你的 `input` 應該有一個 `placeholder` 屬性，值為 `"e.g., John Doe"`。

```js
assert.strictEqual(document.querySelector('fieldset input')?.getAttribute('placeholder'), 'e.g., John Doe');
```

你的 `input` 應該有 `required` 屬性。

```js
assert.isNotNull(document.querySelector('fieldset input[required]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
