---
id: 66a95d0eff8168747805f1f3
title: 第 19 步
challengeType: 0
dashedName: step-19
---

# --description--

當你想讓使用者從列表中選擇多個選項時，你可以使用核取方塊。 

以下是如何處理與食物選項相關的核取方塊的範例：

```html
<fieldset>
  <legend>Food Options</legend>
  <input type="checkbox" id="pizza" name="food" value="pizza">
  <label for="pizza">Pizza</label>
  <input type="checkbox" id="burger" name="food" value="burger">
  <label for="burger">Burger</label>
</fieldset>
```

`value` 屬性用於指定在提交表單時將傳送到伺服端的值。

在你的 `legend` 元素下方，為你的 `input` 添加一個勾選框，並將 `id` 和 `value` 屬性設為 `"ads"`，將 `name` 屬性設為 `"choice"`。

在你的核取方塊 `input` 下方，添加一個 `label` 元素，文字為 `Social Media Ads`。`for` 屬性應設為 `"ads"`。

# --hints--

你應該在你的 `legend` 元素下方放置一個核取方塊 `input`。

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']"));
```

你的核取方塊 `input` 應該有一個 `id` 屬性設定為 `"ads"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.id, "ads");
```

你的核取方塊 `input` 應該有一個設定為 `"choice"` 的 `name` 屬性。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.name, "choice");
```

你的核取方塊 `input` 應該有一個設定為 `"ads"` 的 `value` 屬性。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.value, "ads");
```

你應該在你的核取方塊 `input` 下方放置一個 `label` 元素。

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label"));
```

你的 `label` 元素應該有文字 `Social Media Ads`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.textContent?.trim(), "Social Media Ads");
```

你的 `label` 元素應該有一個 `for` 屬性設定為 `"ads"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.getAttribute("for"), "ads");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
