---
id: 66a9577022877d72d8f43b4f
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

表格的下一部分將詢問用戶爲什麼選擇該酒店。 用戶將有機會選擇多個選項。

首先添加另一個 `fieldset` 元素。

在 `fieldset` 元素內部，添加一個 `legend` 元素，其文本爲 `Why did you choose to stay at our hotel? (Check all that apply)`。

# --hints--

你應該有一個 `fieldset` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 3);
```

你應該在 `fieldset` 元素中包含一個 `legend` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(3) legend'), 1);
```

你的 `legend` 元素的文本應爲 `Why did you choose to stay at our hotel? (Check all that apply)`。 仔細檢查拼寫錯誤和空格問題。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) legend')?.innerText.trim(), 'Why did you choose to stay at our hotel? (Check all that apply)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        --fcc-editable-region--

        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
