---
id: 66f26c32ec6f90df01a44f60
title: Progetta un modulo di contatto
challengeType: 25
dashedName: design-a-contact-form
demoType: onClick
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un `div` con una classe `form-container` per contenere il modulo.

2. Dovresti avere un elemento `form` all’interno dell’elemento `.form-container`.

3. All’interno dell’elemento `form` dovresti avere i seguenti elementi e campi di input:

   - Un elemento `h2` con del testo.
   - Un campo di input di tipo testo per il nome con il `type` impostato su `text` e gli attributi `id`, `name`, `required`.
   - Un campo di input di tipo email con il `type` impostato su `email` e gli attributi `id`, `name`, `required`.
   - Un’area di testo per il messaggio con `id`, `name` e l’attributo `required`.
   - Un elemento `button` con il testo `"Submit"` e un attributo `type` con valore `submit`.

4. Ogni elemento `input` e `textarea` dovrebbe avere il corrispondente elemento `label` con un attributo `for` impostato sul `id` dell’elemento.

5. L’elemento `.form-container` dovrebbe avere un colore di sfondo e valori per `border-radius`, `padding` e `width`.

6. Gli elementi `label` dovrebbero avere un margine e un colore del carattere tipografico.

7. Gli elementi `input` e `textarea` dovrebbero avere valori per:

   - `width`.
   - `padding`.
   - Margine inferiore.

8. L’elemento `button` dovrebbe avere:

   - Un colore di sfondo
   - Una dimensione del carattere tipografico diversa dal predefinito.

9. Il pulsante dovrebbe avere un effetto al passaggio del mouse che cambia il colore di sfondo.

**Nota:** Assicurati di collegare il foglio di stile nel tuo HTML e di applicare il CSS.

# --hints--

Dovresti avere un `div` con una classe `form-container`.

```js
const el = document.querySelector('div.form-container');
assert.isNotNull(el);
```

Dovresti avere un elemento `form` all’interno dell’elemento `.form-container`.

```js
const el = document.querySelector('.form-container form');
assert.isNotNull(el);
```

All’interno del `form` dovresti avere un elemento `h2` con del testo.

```js
const el = document.querySelector('.form-container form h2');
assert.isNotEmpty(el.innerText);
// assert.isAtLeast(document.querySelector('h1').innerText.length, 1);

```

All’interno del `form` dovresti avere un campo di input di tipo impostato su `text`.

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert.isNotNull(el);
```

Il tuo campo di input `type=text` dovrebbe avere un `id`, un `name` e un attributo `required`.

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert.exists(el);
assert.isTrue(el.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

All’interno del `form` dovresti avere un campo di input di tipo email con il `type` impostato su `email`.

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert.isNotNull(el);
```

Il tuo campo di input `type=email` dovrebbe avere un `id`, un `name` e un attributo `required`.

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert.isTrue(el?.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

All’interno del `form` dovresti avere un `textarea` con un `id`, `name` e un attributo `required`.

```js
const el = document.querySelector('.form-container form textarea');
assert.isNotNull(el);
assert.isNotNull(document.querySelector('.form-container form textarea[required]'));
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

All’interno del `form` dovresti avere un elemento `button` con un attributo `type` di valore `submit`.

```js
const el = document.querySelector('.form-container form button');
assert.strictEqual(el?.getAttribute('type'), 'submit');
```

Il tuo pulsante dovrebbe avere un testo `"Submit"`.

```js
const el = document.querySelector('.form-container form button');
assert.strictEqual(document.querySelector('.form-container form button').textContent, 'Submit');
```

Dovresti avere un elemento `label` per l’elemento `#name` con un attributo `for` che corrisponde al valore `id` sull’elemento `input` corrispondente.

```js
const nameInput = document.querySelector('.form-container form input[type="text"]')?.getAttribute('id');

const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

// check if id has a value
assert.isNotEmpty(nameInput);
let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === nameInput) {
        found = true;
        break;
    }
}

assert.isTrue(found);
```

Dovresti avere un elemento `label` per l’elemento `#email` con un attributo `for` che corrisponde al valore `id` sull’elemento `input` corrispondente.

```js
const emailInput = document.querySelector('.form-container form input[type="email"]')?.getAttribute('id');
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);
// check if id has a value
assert.isNotEmpty(emailInput);
let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === emailInput) {
        found = true;
        break;
    }
}

assert.isTrue(found);
```

Dovresti avere un elemento `label` per l’elemento `#message` con un attributo `for` che corrisponde al valore `id` sull’elemento `textarea`.

```js
const messageInput = document.querySelector('.form-container form textarea')?.getAttribute('id');
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

// check if id has a value

assert.isNotEmpty(messageInput);

assert.exists(document.querySelector(`label[for="${messageInput}"]`));
```

Le etichette non dovrebbero essere vuote.

```js
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

for (let label of labels) {
    const labelText = label.textContent.trim();
    assert.isNotEmpty(labelText);
}

```

L’elemento `.form-container` dovrebbe avere un colore di sfondo.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.backgroundColor);
```

L’elemento `.form-container` dovrebbe avere un `border-radius`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.borderRadius);
```

L’elemento `.form-container` dovrebbe avere una spaziatura interna.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.padding);
```

L’elemento `.form-container` dovrebbe avere una larghezza.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.width);
```

Le tue etichette dovrebbero avere un margine e un colore del carattere tipografico.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('label')?.margin);
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('label')?.color);
```

I tuoi elementi `input` dovrebbero avere una larghezza, una spaziatura interna e un margine inferiore.

```js
// Helper function to extract properties from multiple selectors
function getStyleProperties(selectors, properties) {
  const result = {};
  properties.forEach((property) => {
    result[property] = selectors.map((selector) => selector?.[property]).find(Boolean);
  });
  return result;
}

// Initializing selectors
const selectors = [
  new __helpers.CSSHelp(document).getStyle('input'),
  new __helpers.CSSHelp(document).getStyle('input, textarea'),
  new __helpers.CSSHelp(document).getStyle('textarea, input')
];

// Extract width, padding, and margin properties
const { width } = getStyleProperties(selectors, ['width']);
assert.isNotEmpty(width);

const {
  padding,
  paddingTop,
  paddingBottom,
  paddingLeft,
  paddingRight
} = getStyleProperties(selectors, ['padding', 'paddingTop', 'paddingBottom', 'paddingLeft', 'paddingRight']);

// Assert padding properties
assert.isOk(
  padding ||
  (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

const { marginBottom } = getStyleProperties(selectors, ['marginBottom']);
assert.isNotEmpty(marginBottom);

```

Il tuo elemento `textarea` dovrebbe avere una larghezza, una spaziatura interna e un margine inferiore.

```js
// Helper function to extract properties from multiple selectors
function getStyleProperties(selectors, properties) {
  const result = {};
  properties.forEach((property) => {
    result[property] = selectors.map((selector) => selector?.[property]).find(Boolean);
  });
  return result;
}

// Initializing selectors
const selectors = [
  new __helpers.CSSHelp(document).getStyle('textarea'),
  new __helpers.CSSHelp(document).getStyle('input, textarea'),
  new __helpers.CSSHelp(document).getStyle('textarea, input')
];

// Extract width, padding, and margin properties
const { width } = getStyleProperties(selectors, ['width']);
assert.isNotEmpty(width);

const {
  padding,
  paddingTop,
  paddingBottom,
  paddingLeft,
  paddingRight
} = getStyleProperties(selectors, ['padding', 'paddingTop', 'paddingBottom', 'paddingLeft', 'paddingRight']);

// Assert padding properties
assert.isOk(
  padding ||
  (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

const { marginBottom } = getStyleProperties(selectors, ['marginBottom']);
assert.isNotEmpty(marginBottom);
```

Il tuo elemento `button` dovrebbe avere un colore di sfondo.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button')?.backgroundColor);
```

Il tuo elemento `button` dovrebbe avere una dimensione del carattere tipografico.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button')?.fontSize);
```

Il tuo elemento `button` dovrebbe avere un effetto al passaggio del mouse che cambia il colore di sfondo.

```js
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
const btn_hover = gs('button:hover')?.backgroundColor;
assert.isNotEmpty(btn_hover);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Contact Form</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Styled Form Project</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="form-container">
        <form>
            <h2>Contact Us</h2>
            <label for="name">Name:</label>
            <input type="text" id="name" name="name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" required>

            <label for="message">Message:</label>
            <textarea id="message" name="message" required></textarea>

            <button type="submit">Submit</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background: linear-gradient(45deg, #83a4d4, #b6fbff);
    font-family: Arial, sans-serif;
}

.form-container {
    background-color: white;
    border-radius: 10px;
    padding: 20px;
    width: 300px;
    text-align: center;
}

form h2 {
    margin-bottom: 20px;
    color: #333;
}

label {
    display: block;
    margin: 10px 0 5px;
    color: #555;
}

input{
    width: 100%;
    padding: 10px;
    margin-bottom: 10px;
    border: 1px solid #ddd;
    border-radius: 5px;
    box-sizing: border-box;
    resize: vertical;
}

textarea {
    width: 100%;
    padding: 10px;
    margin-bottom: 10px;
    border: 1px solid #ddd;
    border-radius: 5px;
    box-sizing: border-box;
    resize: vertical;
}

button {
    background-color: #4CAF50;
    color: white;
    padding: 10px 15px;
    border: none;
    border-radius: 5px;
    cursor: pointer;
    font-size: 1em;
}

button:hover {
    background-color: #45a049;
}
```
