---
id: 67172b43f84bcd2dec238a3d
title: Tengeneza mpangaji wa vitabu
challengeType: 26
dashedName: build-a-book-organizer
---

# --description--

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuwa na safu ya data ya vitu vinavyoitwa `books` ambapo kila kitu katika safu hiyo kinapaswa kuwa na mfuatano wa herufi `title`, mfuatano mwingine wa herufi `authorName`, na nambari `releaseYear`.

2. Safu yako ya `books` inapaswa kuwa na vitu angalau vitatu.

3. Unapaswa kuwa na kitendakazi cha mwitiko kinachoitwa `sortByYear` kinachokubali vitabu viwili kama vigezo kwa ajili ya kupanga safu ya data.

4. Kitendakazi cha `sortByYear` kinapaswa kurudisha `-1` ikiwa `releaseYear` ya kitabu cha kwanza ni ndogo kuliko ile ya kitabu cha pili.

5. Kitendakazi cha `sortByYear` kinapaswa kurudisha `1` ikiwa `releaseYear` ya kitabu cha kwanza ni kubwa kuliko ile ya kitabu cha pili.

6. Kitendakazi cha `sortByYear` kinapaswa kurudisha `0` ikiwa thamani za `releaseYear` zote ni sawa.

7. Unapaswa kuchuja vitabu vilivyoandikwa baada ya mwaka fulani kama 1950 kutoka kwenye safu ya `books` na kuhifadhi safu iliyochujwa katika safu mpya iitwayo `filteredBooks`.

8. Unapaswa kupanga vitabu katika safu ya `filteredBooks` kulingana na `releaseYear` kwa mpangilio wa kuongezeka. Umejifunza katika funzo la awali kuwa njia ya `sort()` itapanga safu hiyo mahali pake. Hii inamaanisha safu ya `filteredBooks` itabadilishwa.

# --hints--

Unapaswa kuwa na kitendakazi `sortByYear` katika msimbo wako.

```js
assert.isFunction(sortByYear);
```

Kitendakazi chako cha `sortByYear` kinapaswa kuchukua vigezo viwili.

```js
assert.lengthOf(sortByYear, 2);
```

Kitendakazi chako cha `sortByYear` kinapaswa kurudisha `-1` ikiwa `releaseYear` ya kitu cha `book1` ni ndogo kuliko ile ya kitu cha `book2`, `1` ikiwa `releaseYear` ya kitu cha `book1` ni kubwa kuliko ile ya kitu cha `book2`, na `0` katika hali nyingine zote.

```js
assert.equal(sortByYear({releaseYear: 1913}, {releaseYear: 1925}), -1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1913}), 1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1925}), 0);
```

Unapaswa kuwa na safu ya data `books` katika msimbo wako.

```js
assert.isArray(books)
```

Safu yako ya `books` inapaswa kuwa na vitu vya vitabu vitatu angalau.

```js
assert.isAtLeast(books.length, 3)
```

Safu yako ya `books` inapaswa kuwa na vitu vyenye mfuatano wa herufi `title`, mfuatano mwingine wa herufi `authorName`, na nambari `releaseYear`.

```js

books.forEach((book) => {
  assert.isObject(book);
  assert.sameMembers(Object.keys(book), ["title", "authorName", "releaseYear"]);
  assert.isString(book.title);
  assert.isNotEmpty(book.title);
  assert.isString(book.authorName);
  assert.isNotEmpty(book.authorName);
  assert.isNumber(book.releaseYear);
})
```

Unapaswa kuwa na safu ya data `filteredBooks` katika msimbo wako.

```js
assert.isArray(filteredBooks)
```

Safu ya `filteredBooks` inapaswa kuwa na baadhi ya vitabu vilivyochujwa kutoka safu ya `books` na isiwe tupu.

```js
assert.isBelow(filteredBooks.length, books.length)
assert.isNotEmpty(filteredBooks)
```

Safu ya `filteredBooks` inapaswa kuwa na vitabu tu kutoka safu ya awali ya `books`.

```js
filteredBooks.forEach(book => {
  assert.isTrue(
    books.some(original => 
      original.title === book.title &&
      original.authorName === book.authorName &&
      original.releaseYear === book.releaseYear
    )
  );
});
```

Safu ya `filteredBooks` inapaswa kujumuisha vitabu vilivyotolewa mwaka huo au kabla yake.

```js
const maxYear = Math.max(...filteredBooks.map(book => book.releaseYear));
const removedYears = books.filter(book => !filteredBooks.includes(book)).map(book => book.releaseYear);
assert.isTrue(removedYears.every(year => year > maxYear));
```

Unapaswa kuita kitendakazi cha ngazi ya juu cha `sort` kwa kupitisha kitendakazi cha mwitiko cha `sortByYear` kwenye safu ya `filteredBooks`.

```js
assert.match(__helpers.removeJSComments(code), /\s*filteredBooks\s*\.\s*sort\s*\(\s*sortByYear\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sortByYear(book1, book2) {
  if (book1.releaseYear < book2.releaseYear) {
    return -1;
  }

  if (book1.releaseYear > book2.releaseYear) {
    return 1;
  }

  return 0;
}

const books = [
  {
    title: "The Great Gatsby",
    authorName: "F. Scott Fitzgerald",
    releaseYear: 1925,
  },
  {
    title: "In Search of Lost Time",
    authorName: "Marcel Proust",
    releaseYear: 1913,
  },
  {
    title: "Ulysses",
    authorName: "James Joyce",
    releaseYear: 1922,
  },
  {
    title: "One Hundred Years of Solitude",
    authorName: "Gabriel García Márquez",
    releaseYear: 1967,
  },
  {
    title: "The Catcher in the Rye",
    authorName: "J. D. Salinger",
    releaseYear: 1951,
  },
];

let filteredBooks = books.filter((book) => book.releaseYear < 1950);

filteredBooks.sort(sortByYear);

console.log("Books Written Before 1950 (sorted according to release year)");
filteredBooks.forEach((book) => {
  console.log(`${book.title} by ${book.authorName} (${book.releaseYear})`);
});
```
