---
id: 587d7dbf367417b2b2512bbc
title: Dividi i tuoi stili in parti più piccole con i partials
challengeType: 0
forumTopicId: 301459
dashedName: split-your-styles-into-smaller-chunks-with-partials
---

# --description--

<dfn>Partials</dfn> in Sass sono file separati che contengono segmenti di codice CSS. Questi vengono importati e usati in altri file Sass. È un ottimo modo per raggruppare codice simile in un modulo per mantenerlo organizzato.

I nomi dei partials iniziano con il carattere underscore (`_`), che indica a Sass che si tratta di un piccolo segmento di CSS e che non deve convertirlo in un file CSS. Inoltre, i file Sass terminano con l'estensione `.scss`. Per importare il codice del partial in un altro file Sass, usa la direttiva `@import`.

Ad esempio, se tutti i tuoi mixin sono salvati in un partial chiamato "\_mixins.scss" e ti servono nel file "main.scss", ecco come usarli nel file principale:

```scss
@import 'mixins'
```

Nota che nell'istruzione `import` non serve né l'underscore né l'estensione del file — Sass capisce che si tratta di un partial. Una volta importato un partial in un file, tutte le variabili, i mixin e gli altri codici sono disponibili per l'uso.

# --instructions--

Scrivi un'istruzione `@import` per importare un partial chiamato `_variables.scss` nel file main.scss.

# --hints--

Il tuo codice dovrebbe usare la direttiva `@import` e non includere l'underscore nel nome del file.

```js
assert(code.match(/@import\s+?('|")variables\1/gi));
```

# --seed--

## --seed-contents--

```html
<!-- The main.scss file -->
```

# --solutions--

```html
@import 'variables'
```
