---
id: 587d7dbf367417b2b2512bbb
title: applicare uno stile finché non si verifica una condizione con @while
challengeType: 0
forumTopicId: 301454
dashedName: apply-a-style-until-a-condition-is-met-with-while
---

# --description--

La direttiva `@while` è un'opzione con funzionalità simili al loop `while` di JavaScript. Crea regole CSS finché non si verifica una condizione.

La sfida `@for` ha fornito un esempio per creare un sistema di griglia semplice. Questo può funzionare anche con `@while`.

```scss
$x: 1;
@while $x < 13 {
  .col-#{$x} { width: 100%/12 * $x;}
  $x: $x + 1;
}
```

Per prima cosa, definisci una variabile `$x` e impostala su 1. Poi usa la direttiva `@while` per creare il sistema di griglia *finché* `$x` è minore di 13. Dopo aver impostato la regola CSS per `width`, `$x` viene incrementato di 1 per evitare un loop infinito.

# --instructions--

Usa `@while` per creare una serie di classi con `font-sizes` diversi.

Dovrebbero esserci 5 classi diverse da `text-1` a `text-5`. Poi imposta `font-size` su `15px` moltiplicato per il numero dell'indice corrente. Assicurati di evitare un loop infinito!

# --hints--

Il tuo codice dovrebbe usare la direttiva `@while`.

```js
assert(code.match(/@while /g));
```

Il tuo codice dovrebbe usare una variabile indice che parte da 1.

```js
assert(code.match(/\$.*:\s*?1;/gi));
```

Il tuo codice dovrebbe incrementare la variabile contatore.

```js
assert(code.match(/\$(.*)\s*?:\s*\$\1\s*\+\s*1\s*;/gi));
```

La tua classe `.text-1` dovrebbe avere un `font-size` di `15px`.

```js
assert($('.text-1').css('font-size') == '15px');
```

La tua classe `.text-2` dovrebbe avere un `font-size` di `30px`.

```js
assert($('.text-2').css('font-size') == '30px');
```

La tua classe `.text-3` dovrebbe avere un `font-size` di `45px`.

```js
assert($('.text-3').css('font-size') == '45px');
```

La tua classe `.text-4` dovrebbe avere un `font-size` di `60px`.

```js
assert($('.text-4').css('font-size') == '60px');
```

La tua classe `.text-5` dovrebbe avere un `font-size` di `75px`.

```js
assert($('.text-5').css('font-size') == '75px');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

# --solutions--

```html
<style type='text/scss'>
  $x: 1;
  @while $x < 6 {
    .text-#{$x}{
      font-size: 15px * $x;
    }
    $x: $x + 1;
  }
</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```
