---
id: 587d7dbf367417b2b2512bba
title: Usa @each per iterare sugli elementi di una lista
challengeType: 0
forumTopicId: 301461
dashedName: use-each-to-map-over-items-in-a-list
---

# --description--

L'ultima sfida ha mostrato come la direttiva `@for` usa un valore iniziale e uno finale per iterare un certo numero di volte. Sass offre anche la direttiva `@each` che itera su ogni elemento di una lista o di una mappa. Ad ogni iterazione, la variabile viene assegnata al valore corrente della lista o della mappa.

```scss
@each $color in blue, red, green {
  .#{$color}-text {color: $color;}
}
```

Una mappa ha una sintassi leggermente diversa. Ecco un esempio:

```scss
$colors: (color1: blue, color2: red, color3: green);

@each $key, $color in $colors {
  .#{$color}-text {color: $color;}
}
```

Nota che la variabile `$key` è necessaria per fare riferimento alle chiavi nella mappa. Altrimenti, il CSS compilato conterrebbe `color1`, `color2`... al suo interno. Entrambi gli esempi di codice sopra vengono convertiti nel seguente CSS:

```scss
.blue-text {
  color: blue;
}

.red-text {
  color: red;
}

.green-text {
  color: green;
}
```

# --instructions--

Scrivi una direttiva `@each` che attraversa una lista: `blue, black, red` e assegna ogni variabile a una classe `.color-bg`, dove la parte `color` cambia per ogni elemento con il rispettivo colore. Ogni classe dovrebbe impostare anche la proprietà `background-color` al rispettivo colore.

# --hints--

Il tuo codice dovrebbe usare la direttiva `@each`.

```js
assert(code.match(/@each /g));
```

La tua classe `.blue-bg` dovrebbe avere un `background-color` di blue.

```js
assert($('.blue-bg').css('background-color') == 'rgb(0, 0, 255)');
```

La tua classe `.black-bg` dovrebbe avere un `background-color` di black.

```js
assert($('.black-bg').css('background-color') == 'rgb(0, 0, 0)');
```

La tua classe `.red-bg` dovrebbe avere un `background-color` di red.

```js
assert($('.red-bg').css('background-color') == 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```

# --solutions--

```html
<style type='text/scss'>

  @each $color in blue, black, red {
    .#{$color}-bg {background-color: $color;}
  }

  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```

---

```html
<style type='text/scss'>

  $colors: (color1: blue, color2: black, color3: red);

  @each $key, $color in $colors {
    .#{$color}-bg {background-color: $color;}
  }

  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```
