---
id: 587d7dbd367417b2b2512bb6
title: creare CSS riutilizzabile con i mixin
challengeType: 0
forumTopicId: 301455
dashedName: create-reusable-css-with-mixins
---

# --description--

In Sass, un <dfn>mixin</dfn> è un gruppo di dichiarazioni CSS che possono essere riutilizzate in tutto il foglio di stile. La definizione inizia con la regola at `@mixin`, seguita da un nome personalizzato. Applichi il mixin usando la regola at `@include`.

```scss
@mixin reset-list {
  margin: 0;
  padding: 0;
  list-style: none;
}

nav ul {
  @include reset-list;
}
```

Compila in:

```css
nav ul {
  margin: 0;
  padding: 0;
  list-style: none;
}
```

I tuoi mixin possono anche accettare argomenti, che permettono di personalizzarne il comportamento. Gli argomenti sono obbligatori quando usi il mixin.

```scss
@mixin prose($font-size, $spacing) {
  font-size: $font-size;
  margin: 0;
  margin-block-end: $spacing;
}

p {
  @include prose(1.25rem, 1rem);
}

h2 {
  @include prose(2.4rem, 1.5rem);
}
```

Puoi rendere opzionali gli argomenti assegnando ai parametri valori predefiniti.

```scss
@mixin text-color($color: black) {
  color: $color;
}

p {
  @include text-color(); /* color: black */
}

nav a {
  @include text-color(orange);
}
```

# --instructions--

Scrivi un mixin chiamato `shape` e dagli 3 parametri: `$w`, `$h` e `$bg-color`.

Usa il mixin `shape` per dare all'elemento `#square` una larghezza e un'altezza di `50px` e il colore di sfondo `red`. Per l'elemento `#rect-a` aggiungi una larghezza di `100px`, un'altezza di `50px` e il colore di sfondo `blue`. Infine, per l'elemento `#rect-b` aggiungi una larghezza di `50px`, un'altezza di `100px` e il colore di sfondo `orange`.

# --hints--

Dovresti dichiarare un mixin chiamato `shape` con 3 parametri: `$w`, `$h` e `$bg-color`.

```js
assert.match(code, /@mixin\s+shape\s*\(\s*\$w,\s*\$h,\s*\$bg-color\s*\)\s*{/gi);
```

Il tuo mixin dovrebbe includere una proprietà `width` che usa il parametro `$w`.

```js
assert.match(__helpers.removeWhiteSpace(code), /width:\$w;/gi);
```

Il tuo mixin dovrebbe includere una proprietà `height` che usa il parametro `$h`.

```js
assert.match(__helpers.removeWhiteSpace(code), /height:\$h;/gi);
```

Il tuo mixin dovrebbe includere una proprietà `background-color` che usa il parametro `$bg-color`.

```js
assert.match(__helpers.removeWhiteSpace(code), /background-color:\$bg\-color;/gi);
```

Dovresti sostituire gli stili all'interno del selettore `#square` con una chiamata al mixin `shape` usando la parola chiave `@include`. Impostando una larghezza e un'altezza di `50px` e il colore di sfondo `red`.

```js
assert.match(code, /#square\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*50px\s*,\s*red\s*\)\s*;\s*}/gi);
```

Dovresti sostituire gli stili all'interno del selettore `#rect-a` con una chiamata al mixin `shape` usando la parola chiave `@include`. Impostando una larghezza di `100px`, un'altezza di `50px` e il colore di sfondo `blue`.

```js
assert.match(code, /#rect-a\s*{\s*@include\s+shape\s*\(\s*100px\s*,\s*50px\s*,\s*blue\s*\)\s*;\s*}/gi);
```

Dovresti sostituire gli stili all'interno del selettore `#rect-b` con una chiamata al mixin `shape` usando la parola chiave `@include`. Impostando una larghezza di `50px`, un'altezza di `100px` e il colore di sfondo `orange`.

```js
assert.match(code, /#rect-b\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*100px\s*,\s*orange\s*\)\s*;\s*}/gi);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
#square {
  width: 50px;
  height: 50px;
  background-color: red;
}

#rect-a {
  width: 100px;
  height: 50px;
  background-color: blue;
}

#rect-b {
  width: 50px;
  height: 100px;
  background-color: orange;
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```

# --solutions--

```html
<style type='text/scss'>
@mixin shape($w, $h, $bg-color) {
  width: $w;
  height: $h;
  background-color: $bg-color;
}

#square {
  @include shape(50px, 50px, red);
}

#rect-a {
  @include shape(100px, 50px, blue);
}

#rect-b {
  @include shape(50px, 100px, orange);
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```
