---
id: 587d7dbd367417b2b2512bb5
title: Annidare CSS con Sass
challengeType: 0
forumTopicId: 301457
dashedName: nest-css-with-sass
---

# --description--

Sass permette di annidare le regole CSS, un modo utile per organizzare un foglio di stile.

Normalmente, ogni elemento viene definito su una riga diversa per definirne lo stile, così:

```scss
article {
  height: 200px;
}

article p {
  color: white;
}

article ul {
  color: blue;
}
```

Per un progetto grande, il file CSS avrà molte righe e regole. Qui l’annidamento può aiutare a organizzare il codice inserendo le regole di stile dei figli all’interno dei rispettivi elementi genitori:

```scss
article {
  height: 200px;

  p {
    color: white;
  }

  ul {
    color: blue;
  }
}

```

# --instructions--

Usa la tecnica di annidamento mostrata sopra per riorganizzare le regole CSS per entrambi i figli dell’elemento `.blog-post`. Per i test, `h1` deve venire prima dell’elemento `p`.

# --hints--

Il tuo codice dovrebbe riorganizzare le regole CSS in modo che `h1` e `p` siano annidati nell’elemento genitore `.blog-post`.

```js
assert(
  code.match(
    /\.blog-post\s*?{\s*?h1\s*?{\s*?text-align:\s*?center;\s*?color:\s*?blue;\s*?}\s*?p\s*?{\s*?font-size:\s*?20px;\s*?}\s*?}/gi
  )
);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
  .blog-post {

  }
  h1 {
    text-align: center;
    color: blue;
  }
  p {
    font-size: 20px;
  }
</style>

<div class="blog-post">
  <h1>Blog Title</h1>
  <p>This is a paragraph</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  .blog-post {
    h1 {
      text-align: center;
      color: blue;
    }
    p {
      font-size: 20px;
    }
  }
</style>

<div class="blog-post">
  <h1>Blog Title</h1>
  <p>This is a paragraph</p>
</div>
```
