---
id: 587d7dbd367417b2b2512bb4
title: Archiviare dati con variabili Sass
challengeType: 0
forumTopicId: 301460
dashedName: store-data-with-sass-variables
---

# --description--

Una funzionalità di Sass diversa dal CSS è che usa le variabili. Vengono dichiarate e impostate per archiviare dati, in modo simile a JavaScript.

In JavaScript, le variabili si definiscono usando le parole chiave `let` e `const`. In Sass, le variabili iniziano con un `$` seguito dal nome della variabile.

Ecco un paio di esempi:

```scss
$main-fonts: Arial, sans-serif;
$headings-color: green;
```

E per usare le variabili:

```scss
h1 {
  font-family: $main-fonts;
  color: $headings-color;
}
```

Un esempio in cui le variabili sono utili è quando diversi elementi devono avere lo stesso colore. Se quel colore cambia, l’unico posto dove modificare il codice è il valore della variabile.

# --instructions--

Crea una variabile `$text-color` e impostala su `red`. Poi cambia il valore della proprietà `color` per `.blog-post` e `h2` usando la variabile `$text-color`.

# --hints--

Il tuo codice dovrebbe avere una variabile Sass dichiarata per `$text-color` con valore `red`.

```js
assert(code.match(/\$text-color\s*:\s*?red\s*;/g));
```

Il tuo codice dovrebbe usare la variabile `$text-color` per cambiare la `color` per gli elementi `.blog-post` e `h2`.

```js
assert(code.match(/color\s*:\s*\$text-color\s*;?/g));
```

Il tuo elemento `.blog-post` dovrebbe avere una `color` rossa.

```js
assert($('.blog-post').css('color') == 'rgb(255, 0, 0)');
```

I tuoi elementi `h2` dovrebbero avere una `color` rossa.

```js
assert($('h2').css('color') == 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>


  .header{
    text-align: center;
  }
  .blog-post, h2 {
    color: red;
  }
</style>

<h1 class="header">Learn Sass</h1>
<div class="blog-post">
  <h2>Some random title</h2>
  <p>This is a paragraph with some random text in it</p>
</div>
<div class="blog-post">
  <h2>Header #2</h2>
  <p>Here is some more random text.</p>
</div>
<div class="blog-post">
  <h2>Here is another header</h2>
  <p>Even more random text within a paragraph</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  $text-color: red;

  .header{
    text-align: center;
  }
  .blog-post, h2 {
    color: $text-color;
  }
</style>

<h1 class="header">Learn Sass</h1>
<div class="blog-post">
  <h2>Some random title</h2>
  <p>This is a paragraph with some random text in it</p>
</div>
<div class="blog-post">
  <h2>Header #2</h2>
  <p>Here is some more random text.</p>
</div>
<div class="blog-post">
  <h2>Here is another header</h2>
  <p>Even more random text within a paragraph</p>
</div>
```
