---
id: 63bf511b85b6082e54dc1573
title: Passaggio 24
challengeType: 0
dashedName: step-24
---

# --description--

JavaScript fornisce un metodo `.replace()` che ti permette di sostituire caratteri in una stringa con un'altra stringa. Questo metodo accetta due argomenti. Il primo argomento è la sequenza di caratteri da sostituire, che può essere una stringa o un pattern di espressione regolare. Il secondo argomento è la stringa che sostituisce la sequenza corrispondente.

Poiché le stringhe sono immutabili, il metodo `replace` restituisce una nuova stringa con i caratteri sostituiti.

In questo esempio, il metodo `replace` viene usato per sostituire tutte le occorrenze della lettera `l` con il numero `1` nella stringa `hello`.

```js
"hello".replace(/l/g, "1");
```

Usa la tua `regex` per sostituire tutte le occorrenze di `+`, `-` e uno spazio in `str` con una stringa vuota. Restituisci questo valore.

# --hints--

La tua `cleanInputString` dovrebbe chiamare il metodo `replace` di `str`.

```js
assert.match(cleanInputString.toString(), /str\.replace\(/);
```

Dovresti passare `regex` come primo argomento a `replace`.

```js
assert.match(cleanInputString.toString(), /str\.replace\(\s*regex\s*/);
```

Dovresti passare `""` come secondo argomento a `replace`.

```js
assert.match(cleanInputString.toString(), /str\.replace\(\s*regex\s*,\s*("|')\1\s*\)/);
```

La tua funzione `cleanInputString` dovrebbe restituire direttamente il risultato del metodo `replace`.

```js
assert.match(cleanInputString.toString(), /return\s+str\.replace\(\s*regex\s*,\s*("|')\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
          <div class="controls">
            <span>
              <label for="entry-dropdown">Add food or exercise:</label>
              <select id="entry-dropdown" name="options">
                <option value="breakfast" selected>Breakfast</option>
                <option value="lunch">Lunch</option>
                <option value="dinner">Dinner</option>
                <option value="snacks">Snacks</option>
                <option value="exercise">Exercise</option>
              </select>
              <button type="button" id="add-entry">Add Entry</button>
            </span>
          </div>
          <div>
            <button type="submit">
              Calculate Remaining Calories
            </button>
            <button type="button" id="clear">Clear</button>
          </div>
        </form>
        <div id="output" class="output hide"></div>
      </div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js
const calorieCounter = document.getElementById('calorie-counter');
const budgetNumberInput = document.getElementById('budget');
const entryDropdown = document.getElementById('entry-dropdown');
const addEntryButton = document.getElementById('add-entry');
const clearButton = document.getElementById('clear');
const output = document.getElementById('output');
let isError = false;

--fcc-editable-region--
function cleanInputString(str) {
  const regex = /[+-\s]/g;

}
--fcc-editable-region--
```
