---
id: 5d822fd413a79914d39e98d1
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

Give your `.background-buildings` element a `width` and `height` of `100%` to make it the full width and height of its parent, the `body`.

# --hints--

`background-buildings` クラスを使用して正しい要素を選択する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.background-buildings'));
```

`.background-buildings` 要素の `width` が `100%` でなければなりません。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.width, '100%');
```

`.background-buildings` 要素の `height` が `100%` でなければなりません。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.height, '100%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings"></div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--


--fcc-editable-region--

```

