---
id: 5d822fd413a79914d39e98cd
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

In CSS, you can target everything with an asterisk. Add a border to everything by using the `*` selector, and giving it a `border` of `1px solid black`. This is a trick that helps visualize where elements are and their size. You will remove this later.

# --hints--

`*` セレクターを使用する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

`border` プロパティを使用してすべての要素のスタイルを設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

すべての要素が `1px solid black` の境界線を持つ必要があります。

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
