---
id: 5d822fd413a79914d39e98cb
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

Next, add opening and closing `head` and `body` tags within the `html` element.

# --hints--

`head` の開始タグが 1 つ必要です。

```js
assert(code.match(/<head\s*>/i));
```

`head` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/head\s*>/i));
```

`body` の開始タグが 1 つ必要です。

```js
assert(code.match(/<body\s*>/i));
```

`body` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 要素と `body` 要素は兄弟要素である必要があります。

```js
assert(document.querySelector('head').nextElementSibling.localName === 'body');
```

`head` 要素は `html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'head'));
```

`body` 要素は `html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html lang="en">

</html>
--fcc-editable-region--

```

