---
id: 5d822fd413a79914d39e98cc
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

Within the `head`, nest a `meta` element with a `charset` of `UTF-8`, a `title` element with a title of `City Skyline`, and a `link` element that links your `styles.css` file.

# --hints--

`head` 要素の中に新しい `meta` 要素を作成する必要があります。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('meta'));
```

`meta` タグに、値が `UTF-8` に設定されている `charset` 属性を加える必要があります。

```js
const linkElement = document.querySelector('meta');
const charset = linkElement?.getAttribute("charset").toLowerCase();
assert.strictEqual(charset, 'utf-8');
```

コードには `title` 要素が 1 つ必要です。

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title` 要素は `head` 要素内にある必要があります。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('title'));
```

プロジェクトには `City Skyline` というタイトルが必要です。

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'city skyline')
```

タイトルについて、大文字小文字の区別とスペルに気をつけましょう。

```js
const title = document.querySelector('title');
assert.equal(title.text, 'City Skyline');
```

コードには `link` 要素が 1 つ必要です。

```js
assert.isNotNull(document.querySelector('link'));
```

既存の `head` タグを変更しないようにしてください。 終了タグを削除していないか確認してください。

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

空要素である `link` 要素が 1 つだけ必要です。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

`link` 要素は `head` 要素の中にある必要があります。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

`link` 要素には `rel` 属性があり、値が `stylesheet` に設定されている必要があります。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link` 要素には `href` 属性があり、値が `styles.css` に設定されている必要があります。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
