---
id: 5d822fd413a79914d39e98d4
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

Give the parts of your building `width` and `height` properties with these values: `70%` and `10%` to `.bb1a`, `80%` and `10%` to `.bb1b`, `90%` and `10%` to `.bb1c`, and `100%` and `70%` to `.bb1d`. Remember that these percentages are relative to the parent and note that the heights will add up to 100% - vertically filling the container.

# --hints--

クラスセレクターを使用して `.bb1a` のスタイルを指定してください。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a'));
```

`.bb1a` の `width` を `70%` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.width, '70%');
```

`.bb1a` の `height` を `10%` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.height, '10%');
```

クラスセレクターを使用して `.bb1b` のスタイルを指定してください。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b'));
```

`.bb1b` の `width` を `80%` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.width, '80%');
```

`.bb1b` の `height` を `10%` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.height, '10%');
```

クラスセレクターを使用して `.bb1c` のスタイルを指定してください。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c'));
```

`.bb1c` の `width` を `90%` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.width, '90%');
```

`.bb1c` の `height` を `10%` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.height, '10%');
```

クラスセレクターを使用して `.bb1d` のスタイルを指定してください。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d'));
```

`.bb1d` の `width` を `100%` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.width, '100%');
```

`.bb1d` の `height` を `70%` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
--fcc-editable-region--

--fcc-editable-region--

```

