---
id: 66c644c08ebcd2ef90c750bd
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Nos vídeos da aula anterior, você aprendeu como adicionar itens ao final de um array usando o método `push` assim:

```js
fruits.push('pear');
```

Usando o método `push`, adicione a string `"Apples"` ao array `shoppingList`.

# --hints--

O array `shoppingList` deve conter um item.

```js
assert.lengthOf(shoppingList, 1);
```

O array `shoppingList` deve conter a string `"Apples"`. Verifique novamente se há erros de grafia e capitalização.

```js
assert.includeMembers(shoppingList, ['Apples']);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

--fcc-editable-region--

--fcc-editable-region--
```
