---
id: 66c64095f61166ede6e39a84
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Para a lista de compras, você usará um array para representar os itens que precisa comprar.

In the previous lectures, you learned how to create arrays like this:

```js
const fruits = ['apple', 'banana', 'orange'];
```

Neste passo, crie uma variável chamada `shoppingList` e atribua a ela um array vazio.

# --hints--

Você deve ter uma variável chamada `shoppingList`.

```js
assert.isNotNull(shoppingList);
```

Você deve atribuir um array à variável `shoppingList`.

```js
assert.isArray(shoppingList);
```

O array deve estar vazio.

```js
assert.isEmpty(shoppingList);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

--fcc-editable-region--

--fcc-editable-region--
```
