---
id: 66c64ad16796c7f2419b45c5
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

Agora é hora de começar a adicionar itens ao topo da lista de compras.

Comece adicionando um `console.log` que registre a mensagem `"It looks like we need to get some cooking oil."` no console.

# --hints--

Você deve registrar a mensagem `"It looks like we need to get some cooking oil."` no console.

```js
assert.match(code, /console\.log\((['"])(It\s+looks\s+like\s+we\s+need\s+to\s+get\s+some\s+cooking\s+oil\.)\1\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg() {
  return `Current Shopping List: ${shoppingList}`;
}

console.log(getShoppingListMsg());

shoppingList.push("Grapes");
console.log(getShoppingListMsg());

--fcc-editable-region--

--fcc-editable-region--
```
