---
id: 66c645b345a39bf04864dc50
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Para ver a lista de compras atualizada, você precisará registrar a lista de compras atual e uma mensagem curta no console.

Como a mensagem será usada repetidamente ao longo do workshop, é melhor criar uma função reutilizável.

Create a function called `getShoppingListMsg` that takes an array as a parameter and returns the string `"Current Shopping List: "` followed by the contents of the provided array.

Você pode usar template literals ou concatenação de strings com o operador `+` para combinar a string e o array `shoppingList`.

# --hints--

Você deve ter uma função chamada `getShoppingListMsg` no código.

```js
assert.isFunction(getShoppingListMsg);
```

A função `getShoppingListMsg` deve retornar uma string.

```js
assert.isString(getShoppingListMsg(shoppingList));
```

Your `getShoppingListMsg` function should accept an array as an argument.


```js
assert.strictEqual(getShoppingListMsg(["Anchovies", "Olives"]), "Current Shopping List: Anchovies,Olives");
```

Your `getShoppingListMsg` function should accept an array and return the message `"Current Shopping List: "` followed by the contents of that array. Verifique novamente os erros de espaçamento e ortografia.

```js
assert.strictEqual(getShoppingListMsg(["Apples"]), "Current Shopping List: Apples");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

--fcc-editable-region--

--fcc-editable-region--
```
