---
id: 66c73a7798f6f62b2ae58f22
title: Passo 16
challengeType: 1
dashedName: step-16
---

# --description--

Agora é hora de adicionar mais itens ao começo da lista de compras.

Comece adicionando uma instrução `console.log` que registre a mensagem `"It might be nice to get a dessert."`

Abaixo dessa instrução `console`, use o método de array correto para adicionar a string `"Chocolate Cake"` ao começo do array `shoppingList`.

Finalmente, adicione um `console.log` e chame a função `getShoppingListMsg` dentro do `console.log` para ver a lista atualizada registrada no console.

# --hints--

Você deve ter uma instrução `console.log` que registre a mensagem `"It might be nice to get a dessert."`.

```js
assert.match(code, /console\.log\((['"])(It\s+might\s+be\s+nice\s+to\s+get\s+a\s+dessert\.)\1\)/);
```

O array `shoppingList` deve ter um total de `6` itens.

```js
assert.lengthOf(shoppingList, 6);
```

Você deve usar o método `unshift` para adicionar a string `"Chocolate Cake"` ao começo do array `shoppingList`.

```js
assert.strictEqual(shoppingList[0], "Chocolate Cake");
```

Você deve chamar a função `getShoppingListMsg` dentro do `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\)\s*\)/g), 6);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg() {
  return `Current Shopping List: ${shoppingList}`;
}

console.log(getShoppingListMsg());

shoppingList.push("Grapes");
console.log(getShoppingListMsg());

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg());

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg());

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg());

--fcc-editable-region--

--fcc-editable-region--
```
