---
id: 66c73a0c5b264f2a75164d94
title: Passo 15
challengeType: 1
dashedName: step-15
---

# --description--

Agora é hora de registrar o array `shoppingList` atualizado no console.

Adicione um `console.log` e chame a função `getShoppingListMsg` dentro do `console.log` para ver a lista atualizada registrada no console.

# --hints--

Você deve chamar a função `getShoppingListMsg` dentro do `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\)\s*\)/g), 5);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg() {
  return `Current Shopping List: ${shoppingList}`;
}

console.log(getShoppingListMsg());

shoppingList.push("Grapes");
console.log(getShoppingListMsg());

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg());

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg());

console.log("This looks like too much junk food.");

shoppingList.pop();

--fcc-editable-region--

--fcc-editable-region--
```
