---
id: 66c73fa7433e082c4be096b1
title: Passo 17
challengeType: 1
dashedName: step-17
---

# --description--

Na última parte do workshop, você revisará como remover um item do início de um array.

Comece inserindo um `console.log` que registre a mensagem `"On second thought, maybe we should be more health conscious."`.

# --hints--

Você deve ter um `console.log` que registre a mensagem `"On second thought, maybe we should be more health conscious."`.

```js
assert.match(code, /console\.log\((["'])(On\s+second\s+thought,\s+maybe\s+we\s+should\s+be\s+more\s+health\s+conscious\.)\1\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg() {
  return `Current Shopping List: ${shoppingList}`;
}

console.log(getShoppingListMsg());

shoppingList.push("Grapes");
console.log(getShoppingListMsg());

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg());

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg());

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg());

console.log("It might be nice to get a dessert.");

shoppingList.unshift("Chocolate Cake");
console.log(getShoppingListMsg());

--fcc-editable-region--

--fcc-editable-region--
```
