---
id: 66c72b0ffbc5522525768558
title: Passo 11
challengeType: 1
dashedName: step-11
---

# --description--

Nos passos anteriores, você revisou como adicionar um item ao final do array usando o método `push`.

O método `push`, no entanto, aceita múltiplos argumentos. Então, você pode adicionar múltiplos itens ao final do array assim:

```js
array.push(item1, item2, item3);
```

Neste passo, use o método `push` para adicionar as strings `"Popcorn", "Beef Jerky", "Potato Chips"` ao array `shoppingList`.

A ordem é importante, então certifique-se de adicionar os itens na ordem listada.

# --hints--

O array `shoppingList` deve ter um total de `6` strings.

```js
assert.lengthOf(shoppingList, 6);
```

O `shoppingList` deve ter a string `"Popcorn"`.

```js
assert.strictEqual(shoppingList[3], "Popcorn");
```

O `shoppingList` deve ter a string `"Beef Jerky"`.

```js
assert.strictEqual(shoppingList[4], "Beef Jerky");
```

O `shoppingList` deve ter a string `"Potato Chips"`.

```js
assert.strictEqual(shoppingList[5], "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg() {
  return `Current Shopping List: ${shoppingList}`;
}

console.log(getShoppingListMsg());

shoppingList.push("Grapes");
console.log(getShoppingListMsg());

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg());

--fcc-editable-region--

--fcc-editable-region--
```
