---
id: 66c742d045c9fc2e09fa64b1
title: Passo 19
challengeType: 1
dashedName: step-19
---

# --description--

A última alteração a fazer na lista de compras é atualizar o primeiro item na lista.

Nos vídeos da aula anterior, você aprendeu como atualizar um item usando notação de colchetes e o índice do item que deseja atualizar.

Aqui está um lembrete de como atualizar um item em um array:

```js
const array = [1, 2, 3, 4, 5];

array[0] = 10;
// Result: [10, 2, 3, 4, 5]
console.log(array); 
```

Atualize o primeiro item no array `shoppingList` para que seja `"Canola Oil"`.

# --hints--

Você deve atualizar o primeiro item no array `shoppingList` para `"Canola Oil"`.

```js
assert.equal(shoppingList[0], "Canola Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg() {
  return `Current Shopping List: ${shoppingList}`;
}

console.log(getShoppingListMsg());

shoppingList.push("Grapes");
console.log(getShoppingListMsg());

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg());

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg());

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg());

console.log("It might be nice to get a dessert.");

shoppingList.unshift("Chocolate Cake");
console.log(getShoppingListMsg());

console.log("On second thought, maybe we should be more health conscious.");

shoppingList.shift();

--fcc-editable-region--

--fcc-editable-region--
```
