---
id: 66c730ee6ae076281721d0b9
title: Passo 14
challengeType: 1
dashedName: step-14
---

# --description--

In the previous lectures, you learned how to remove items from the end of an array using the `pop` method.

Aqui está um lembrete de como usar o método `pop`:

```js
let array = [1, 2, 3, 4, 5];
array.pop();

// [1, 2, 3, 4]
console.log(array); 
```

Use o método `pop` para remover o último item do array `shoppingList`.

# --hints--

O array `shoppingList` agora deve ter `5` itens nele.

```js
assert.lengthOf(shoppingList, 5);
```

O array `shoppingList` não deve mais ter a string `"Potato Chips"` nele.

```js
assert.notInclude(shoppingList, "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

--fcc-editable-region--

--fcc-editable-region--
```
