---
id: 66c74079c30b1c2d166cb9a4
title: Passo 18
challengeType: 1
dashedName: step-18
---

# --description--

Nos vídeos da aula, você aprendeu como remover um item do início do array usando o método `shift`.

Aqui está um lembrete de como usar o método `shift`:

```js
const array = [1, 2, 3, 4, 5];
array.shift();

// Result: [2, 3, 4, 5]
console.log(array); 
```

Use o método `shift` para remover o primeiro item do array `shoppingList`.

# --hints--

O array `shoppingList` deve ter um total de `5` itens.

```js
assert.lengthOf(shoppingList, 5);
```

Você deve usar o método `shift` para remover o primeiro item do array `shoppingList`.

```js
assert.strictEqual(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

console.log("It might be nice to get a dessert.");

shoppingList.unshift("Chocolate Cake");
console.log(getShoppingListMsg(shoppingList));

console.log("On second thought, maybe we should be more health conscious.");

--fcc-editable-region--

--fcc-editable-region--
```
