---
id: 66c72a55418cc9247b710827
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

Em seguida, adicione um `console.log` e chame a função `getShoppingListMsg` dentro do `console.log` para ver a lista atualizada registrada no console.

# --hints--

Você deve chamar a função `getShoppingListMsg` dentro do `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\)\s*\)/g), 3);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(){
  return `Current Shopping List: ${shoppingList}`;
}

console.log(getShoppingListMsg());

shoppingList.push("Grapes");
console.log(getShoppingListMsg());

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");

--fcc-editable-region--

--fcc-editable-region--
```
