---
id: 66cbe2319d3845545a293a0b
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Agora é hora de ver a mensagem registrada no console.

Adicione um `console.log` e chame a função `getShoppingListMsg` dentro do `console.log` para ver a mensagem registrada no console.

# --hints--

Você deve chamar a função `getShoppingListMsg` dentro do `console.log`.

```js
assert.match(code, /console\.log\(\s*getShoppingListMsg\(\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg() {
  return `Current Shopping List: ${shoppingList}`;
}

--fcc-editable-region--

--fcc-editable-region--
```
