---
id: 642db8c409d9991d0b3b2f0d
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

你的專案從一個基本的 HTML 容器和一些對應的 CSS 開始。你的第一個任務將是以程式方式產生你的試算表的儲存格。

全域的 `window` 物件表述瀏覽器視窗（或分頁）。它有一個 `onload` 屬性，允許你定義當視窗已載入整個頁面（包括樣式表和腳本）時的行為。

首先將 `window` 的 `onload` 屬性指定為一個沒有參數的箭頭函式（程式）。在該函式中，宣告一個 `container` 變數，並將透過 `id` 為 `"container"` 取得元素的值指定給它。

# --hints--

你應該存取 `window` 物件的 `onload` 屬性。

```js
assert.match(code, /window\.onload/);
```

你應該將 `onload` 屬性設定為一個函式（程式）。

```js
assert.isFunction(window.onload);
```

你應該使用箭頭語法。

```js
assert.match(code, /window\.onload\s*=\s*\(/);
```

你的 `onload` 函式（程式）不應該帶有任何參數。

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)/);
```

你應該在你的 `onload` 函式中宣告一個 `container` 變數。

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*(?:let|var|const)\s+container/);
```

你的 `container` 變數應該用 `const` 宣告。

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container/);
```

你應該使用 `document.getElementById()`

```js
assert.match(code, /document\.getElementById\(/);
```

你應該取得 `id` 為 `"container"` 的元素。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

你應該將 `#container` 元素指定給 `container`。

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
