---
id: 642ddfdea4200e313f80a4b6
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

請記住，`document` 對象有一個 `.createElement()` 方法，它允許你動態創建新的 HTML 元素。

在 `createLabel` 函數中，聲明一個 `label` 變量併爲其分配一個新的 `div` 元素。

# --hints--

你應該在 `createLabel` 函數中聲明一個 `label` 變量。

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*(?:const|let|var)\s+label/);
```

你的 `label` 變量應使用 `const` 聲明。

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*const\s+label/);
```

你應該使用 `document` 對象的 `.createElement()` 方法。

```js
assert.match(code, /document\.createElement\(/);
```

你應該將字符串 `"div"` 傳遞給 `.createElement()` 方法。

```js
assert.match(code, /document\.createElement\(\s*('|"|`)div\1\s*\)/);
```

你應該將新的 `div` 元素分配給 `label`。

```js
assert.match(code, /const\s+label\s*=\s*document\.createElement\(\s*('|"|`)div\1\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {

  }
}
--fcc-editable-region--
```
