---
id: 642df9df4b5216350de7b0d2
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

你將需要一個函式（程式）來產生一個數字範圍。

宣告一個空的 `range` 函式，該函式接受 `start` 和 `end` 參數。使用 `Array()` 建構式並隱式傳回一個空的陣列。

# --hints--

你應該宣告一個 `range` 變數。

```js
assert.match(code, /(?:let|var|const)\s+range/);
```

你應該使用 `const` 來宣告你的 `range` 變數。

```js
assert.match(code, /const\s+range/);
```

你的 `range` 變數應該是一個函式（程式）。

```js
assert.isFunction(range);
```

你的 `range` 函式應該使用箭頭語法。

```js
assert.match(code, /const\s+range\s*=\s*\(.*\)\s*=>/);
```

你的 `range` 函式應該先接受一個 `start` 參數。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start/)
```

你的 `range` 函式（程式）應該接受第二個 `end` 參數。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

你的 `range` 函式應該使用隱式傳回。請記得這表示你不會使用大括號。

```js
assert.notMatch(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

你的 `range` 函式（程式）應該使用 `Array()` 建構式。原始建構式不需要 `new` 關鍵字。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(/);
```

你不應該傳遞任何東西給 `Array()` 建構式。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
