---
id: 642e02be7845f13b014cd2b0
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

既然你有了 `range` 函式（程式），你也可以用它來創建一個字母的範圍。

使用 `const` 和箭頭語法宣告一個 `charRange` 函式（程式）。它應該接受 `start` 和 `end` 參數。該函式應該隱式地傳回以 `start` 和 `end` 作為引數呼叫 `range()` 的結果。

# --hints--

你應該宣告一個 `charRange` 變數。

```js
assert.match(code, /(?:let|const|var)\s+charRange/);
```

你的 `charRange` 變數應該使用 `const` 宣告。

```js
assert.match(code, /const\s+charRange/);
```

你的 `charRange` 變數應該是一個函式（程式）。

```js
assert.isFunction(charRange);
```

你的 `charRange` 函式應該使用箭頭語法。

```js
assert.match(code, /const\s+charRange\s*=\s*\(.*\)\s*=>/);
```

你的 `charRange` 函式（程式）應該將 `start` 作為第一個參數。

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start/);
```

你的 `charRange` 函式（程式）應該將 `end` 作為第二個參數。

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

你的 `charRange` 函式應該使用隱式傳回。

```js
assert.notMatch(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

你的 `charRange` 函式應該呼叫你的 `range` 函式。

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(/);
```

你應該將 `start` 和 `end` 作為引數傳遞給你的 `range` 呼叫。

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\s*,\s*end\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
