---
id: 643475e13dc727231acd0f72
title: 第 17 步
challengeType: 0
dashedName: step-17
---

# --description--

現在在你的巢狀 `.forEach()` 呼叫中，宣告一個 `input` 變數。使用 `document` 物件的 `.createElement()` 方法來創建一個 `input` 元素。將 `type` 屬性設為 `"text"`，並將 `id` 屬性設為 `letter + number`。

# --hints--

你應該宣告一個 `input` 變數。

```js
assert.match(code, /(?:var|let|const)\s+input/)
```

你應該使用 `const` 來宣告你的 `input` 變數。

```js
assert.match(code, /const\s+input/)
```

你應該呼叫 `document` 物件的 `.createElement()` 方法。

```js
assert.lengthOf(code.match(/document\.createElement\(/g), 2)
```

你應該將字串 `"input"` 傳遞給 `.createElement()` 方法。

```js
assert.match(code, /document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

你應該將你的新 `input` 元素指定到你的 `input` 變數。

```js
assert.match(code, /const\s+input\s*=\s*document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

你應該存取 `input` 元素的 `type` 屬性。

```js
assert.match(code, /input\.type/);
```

你應該將你的 `input` 元素的 `type` 屬性設為 `"text"`。

```js
assert.match(code, /input\.type\s*=\s*('|"|`)text\1/)
```

你應該存取 `input` 元素的 `id` 屬性。

```js
assert.match(code, /input\.id/);
```

你應該將你的 `input` 元素的 `id` 屬性設為 `letter + number`。

```js
assert.match(code, /input\.id\s*=\s*letter\s\+\snumber/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
--fcc-editable-region--
    letters.forEach(letter => {

    })
--fcc-editable-region--
  })
}
```
