---
id: 6434750c53db16218f41e6e1
title: 第 15 步
challengeType: 0
dashedName: step-15
---

# --description--

請記住 `range()` 傳回一個陣列，所以你可以直接將陣列方法串鏈到函式呼叫。

使用 `1` 和 `99` 作為引數呼叫 `range()`，並串鏈 `.forEach()` 方法。將一個以 `number` 為參數的空回呼傳遞給 `.forEach()` 方法。

# --hints--

你應該呼叫你的 `range()` 函式（程式）。

```js
assert.lengthOf(code.match(/range\(/g), 2);
```

你應該將 `1` 作為第一個引數傳遞給你的 `range()` 呼叫。

```js
assert.match(code, /range\(\s*1/);
```

你應該將 `99` 作為第二個引數傳遞給你的 `range()` 呼叫。

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)/);
```

你應該將 `.forEach()` 方法串鏈到你的 `range()` 呼叫。

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(/);
```

你應該使用箭頭語法將回呼函式（程式）傳遞給 `.forEach()`。

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

你的回呼函式（程式）應該只有 `number` 作為唯一參數。

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\(\s*number\s*\)|number)\s*=>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
--fcc-editable-region--

--fcc-editable-region--
}
```
