---
id: 6434759f78ec812264ff8f34
title: 第 16 步
challengeType: 0
dashedName: step-16
---

# --description--

在你的回呼中，你需要進行兩次函式呼叫。首先呼叫 `createLabel()` 並將 `number` 作為引數傳入。你應該會看到一些數字出現在你的試算表中。

然後在你的 `letters` 陣列上呼叫 `.forEach()` 方法。傳入一個空的回呼函式，該函式帶有一個 `letter` 參數。

# --hints--

你應該呼叫你的 `createLabel()` 函式（程式）。

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\(\s*number\s*\)|number)\s*=>\s*\{\s*createLabel\(/);
```

你應該將 `number` 傳遞給你的 `createLabel()` 呼叫。

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\(\s*number\s*\)|number)\s*=>\s*\{\s*createLabel\(/)
```

你應該在你的 `letters` 陣列上呼叫 `.forEach()` 方法。

```js
assert.lengthOf(code.match(/letters\.forEach\(/g), 2)
```

你應該將帶有箭頭語法的回呼函式傳遞給你的 `.forEach()` 方法。

```js
assert.match(code, /letters\.forEach\(\s*(\([^)]*\)|[^\s()]+)\s*=>\s*\{/)
```

你的回呼函式（程式）應該有一個 `letter` 參數。

```js
assert.match(code, /letters\.forEach\(\s*(\(\s*letter\s*\)|letter)\s*=>\s*\{/)
```

你的回呼函式應該是空的。

```js
assert.match(code, /letters\.forEach\(\s*(\(\s*letter\s*\)|letter)\s*=>\s*\{\s*\}/)
```

你的 `letters.forEach()` 回呼函式（程式）應該巢狀在 `range(1, 99).forEach(number => {}` 回呼函式（程式）內。

```js
assert.match(code, /range\s*\(\s*1\s*,\s*99\s*\)\s*.forEach\s*\(\s*(\(\s*number\s*\)|number)\s*=>\s*\{\s*[^}]*letters\.forEach\(\s*(\(\s*letter\s*\)|letter)\s*=>\s*\{\s*\}\s*\)\s*;?\s*\}\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
--fcc-editable-region--
  range(1, 99).forEach(number => {

  })
--fcc-editable-region--
}
```
