---
id: 6437133052eaf04d7300e622
title: 第 22 步
challengeType: 0
dashedName: step-22
---

# --description--

宣告一個 `average` 函式（程式），它以數字陣列作為 `nums` 參數。它應該傳回陣列中所有數字的平均值。

平均值可以透過將陣列中所有數字的總和除以陣列的長度來計算。請記得你有一個可以使用的 `sum` 函式。

# --hints--

你應該宣告一個 `average` 變數。

```js
assert.match(code, /(?:let|const|var)\s+average/);
```

你應該使用 `const` 來宣告你的 `average` 變數。

```js
assert.match(code, /const\s+average/);
```

你的 `average` 變數應該是一個函式（程式）。

```js
assert.isFunction(average);
```

你的 `average` 函式（程式）應該使用箭頭語法。

```js
assert.match(code, /const\s+average\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

你的 `average` 函式應該有一個 `nums` 參數。

```js
assert.match(code, /const\s+average\s*=\s*(\(\s*nums\s*\)|nums)/);
```

你的 `average` 函式應該使用隱式傳回。

```js
assert.notMatch(code, /const\s+average\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*{/);
```

你的 `average` 函式應該傳回 `nums` 陣列的平均值。

```js
assert.equal(average([1,2,3]), 2);
assert.equal(average([1,2,3,4,5]), 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
