---
id: 6437124c4c03dd4c8fb35d56
title: 第 21 步
challengeType: 0
dashedName: step-21
---

# --description--

宣告一個 `isEven` 函式，該函式接受一個 `num` 參數，並在數字為偶數時傳回 `true`，否則傳回 `false`。使用取餘運算子 `%` 來判斷數字是偶數還是奇數。

# --hints--

你應該宣告一個 `isEven` 變數。

```js
assert.match(code, /(?:let|const|var)\s+isEven/);
```

你應該使用 `const` 來宣告你的 `isEven` 變數。

```js
assert.match(code, /const\s+isEven/);
```

你的 `isEven` 變數應該是一個函式（程式）。

```js
assert.isFunction(isEven);
```

你的 `isEven` 函式（程式）應該使用箭頭語法。

```js
assert.match(code, /const\s+isEven\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

你的 `isEven` 函式（程式）應該有一個 `num` 參數。

```js
assert.match(code, /const\s+isEven\s*=\s*(\(\s*num\s*\)|num)\s*=>/);
```

你的 `isEven` 函式（程式）應該使用取模運算子 `%`。

```js
assert.match(isEven.toString(),  /%/);
```

你的 `isEven` 函式應該為偶數傳回 `true`。

```js
assert.isTrue(isEven(2));
assert.isTrue(isEven(1000));
assert.isTrue(isEven(42));
```

你的 `isEven` 函式（程式）應該為奇數傳回 `false`。

```js
assert.isFalse(isEven(1));
assert.isFalse(isEven(333));
assert.isFalse(isEven(777777777));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
const sum = nums => nums.reduce((acc, el) => acc + el, 0);

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
