---
id: 643715013330824ecaa70442
title: 第 23 步
challengeType: 0
dashedName: step-23
---

# --description--

你的下一個函式（程式）將計算數字陣列的中位數值。首先宣告一個帶有 `nums` 參數的 `median` 箭頭函式（程式）。

在函式中，宣告一個 `sorted` 變數並將排序後的 `nums` 陣列複本設值給它。

你應該使用 `slice()` 方法來建立陣列的淺層複本。

# --hints--

你應該宣告一個 `median` 變數。

```js
assert.match(code, /(?:let|const|var)\s+median/);
```

你應該使用 `const` 來宣告你的 `median` 變數。

```js
assert.match(code, /const\s+median/);
```

你的 `median` 變數應該是一個函式（程式）。

```js
assert.isFunction(median);
```

你的 `median` 函式應該使用箭頭語法。

```js
assert.match(code, /const\s+median\s*=\s*\(?/);
```

你的 `median` 函式應該有一個 `nums` 參數。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)/);
```

你的 `median` 函式不應該使用隱式傳回。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{/);
```

你的 `median` 函式應該有一個 `sorted` 變數。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*(?:let|var|const)\s+sorted/);
```

你應該使用 `const` 來宣告你的 `sorted` 變數。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted/);
```

你應該使用 `.slice()` 將 `nums` 陣列的複本指定給 `sorted`。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)/);
```

你應該將 `.sort()` 方法串鏈到你的 `.slice()` 方法。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(/);
```

你應該將回呼函式傳遞給你的 `sort` 方法，以準確地將數字依照遞增順序排序。為了清晰起見，使用隱喻的傳回。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(\s*\(\s*a\s*,\s*b\s*\)\s*=>\s*a\s*-\s*b\s*\)\s*\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
