---
id: 64496df724dd3716a71fe971
title: 步驟 26
challengeType: 0
dashedName: step-26
---

# --description--

物件屬性由鍵/值配對組成。當宣告物件字面值時，你可以使用簡寫屬性名稱。當使用簡寫屬性名稱語法時，變數名稱會成為屬性鍵，而其值則成為屬性值。

以下範例宣告了一個具有屬性 `userId`、`firstName` 和 `loggedIn` 的 `user` 物件。

```js
const userId = 1;
const firstName = "John";
const loggedIn = true;

const user = {
  userId,
  firstName,
  loggedIn,
};

console.log(user); // { userId: 1, firstName: 'John', loggedIn: true }
```

為了追蹤你所有試算表的函式（程式），宣告一個 `spreadsheetFunctions` 物件。使用縮寫語法，將 `sum`、`average` 和 `median` 設為 `spreadsheetFunctions` 物件的屬性。

# --hints--

你應該宣告一個 `spreadsheetFunctions` 變數。

```js
assert.match(code, /(?:const|let|var)\s+spreadsheetFunctions/);
```

你應該使用 `const` 來宣告你的 `spreadsheetFunctions` 變數。

```js
assert.match(code, /const\s+spreadsheetFunctions/);
```

你的 `spreadsheetFunctions` 變數應該是一個物件。

```js
assert.isObject(spreadsheetFunctions);
```

你的 `spreadsheetFunctions` 物件應該有一個 `sum` 屬性。

```js
assert.property(spreadsheetFunctions, "sum");
```

你的 `sum` 屬性應該是你的 `sum` 函式（程式）。

```js
assert.equal(spreadsheetFunctions?.sum, sum);
```

你的 `spreadsheetFunctions` 物件應該有一個 `average` 屬性。

```js
assert.property(spreadsheetFunctions, "average");
```

你的 `average` 屬性應該是你的 `average` 函式（程式）。

```js
assert.equal(spreadsheetFunctions?.average, average);
```

你的 `spreadsheetFunctions` 物件應該有一個 `median` 屬性。

```js
assert.property(spreadsheetFunctions, "median");
```

你的 `median` 屬性應該是你的 `median` 函式（程式）。

```js
assert.equal(spreadsheetFunctions?.median, median);
```

你應該使用解構語法來指定你的屬性。

```js
const objectText = code.replace(/.*const\s+spreadsheetFunctions\s*=\s*\{([^}]*)}.*/s, "$1");
assert.include(objectText, "sum");
assert.include(objectText, "average");
assert.include(objectText, "median");
assert.notInclude(objectText, ":");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

const median = nums => {
  const sorted = nums.slice().sort((a, b) => a - b);
  const length = sorted.length;
  const middle = length / 2 - 1;
  return isEven(length)
    ? average([sorted[middle], sorted[middle + 1]])
    : sorted[Math.ceil(middle)];
}

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
