---
id: 65ae458e23954c3469e0c209
title: 步驟 18
challengeType: 0
dashedName: step-18
---

# --description--

在早期的專案中你學習過 `setAttribute` 方法。更新 JavaScript 中屬性的另一種方式是使用以下語法：

```js
el.attribute = value;
```

連字 HTML 屬性值（例如 `aria-label`）的屬性名稱遵循駝峰式大小寫，變成 `ariaLabel`。

```js
el.ariaLabel = "Aria Label Value";
```

將 `input` 元素的 `aria-label` 屬性設為與 `id` 屬性相同的值。


# --hints--

你應該有一個 `input.ariaLabel`。

```js
assert.match(code, /input\.ariaLabel\s*=/);
```


你應該給 `input.ariaLabel` 賦予與 `input.id` 相同的值。

```js
assert.match(code, /input\.ariaLabel\s*=\s*((letter\s*\+\s*number)|(input.id))/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
     createLabel(number);
--fcc-editable-region--
    letters.forEach(letter => {
     const input = document.createElement("input");
     input.type = "text";
     input.id = letter + number;
 
    })   
--fcc-editable-region--
  })
}
```
