---
id: 66ea7e7399244c6d4b6052a1
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Der nächste Teil des Workshops überprüft, wie die `Math.random()`-Methode funktioniert.

Beginnen Sie damit, eine weitere `console.log()` hinzuzufügen, die den String `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."` in der Konsole ausgibt.
 
# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten eine `console.log()` haben, die den String `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."` ausgibt.

```js
assert.equal(
  getLogs()[1],
  "The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."
);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);
--fcc-editable-region--

--fcc-editable-region--
```
