---
id: 66ea7fc048a6e26e33c598fe
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Jetzt ist es Zeit, eine Zufallszahl mit der `Math.random()`-Methode zu erzeugen.

Erstellen Sie eine Variable namens `randomNum` und weisen Sie ihr den Wert des Ergebnisses der Ausführung der `Math.random()`-Methode zu.

Geben Sie dann die `randomNum`-Variable in der Konsole aus.

Versuchen Sie, im Code ein Leerzeichen hinzuzufügen, um den Bot neu zu starten und verschiedene generierte Zufallszahlen zu sehen.

# --hints--

Sie sollten eine Variable namens `randomNum` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.randomNum);
```

Sie sollten das Ergebnis des Aufrufs der `Math.random()`-Methode der `randomNum`-Variable zuweisen.

```js
const explorer = await __helpers.Explorer(code);
const randomNum = explorer.variables.randomNum;
assert.equal(randomNum?.value?.toString(), "Math.random()");
```

Sie sollten die `randomNum`-Variable in der Konsole ausgeben.

```js
assert.match(code, /console\.log\(\s*randomNum\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");
--fcc-editable-region--

--fcc-editable-region--
```
