---
id: 66ea8239b9eeb76f12299030
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Der nächste Abschnitt des Workshops dient dazu, zu überprüfen, wie man eine Zufallszahl zwischen zwei Werten generiert.

Beginnen Sie damit, einen weiteren `console.log()` hinzuzufügen, der die Nachricht `"Now, generate a random number between two values."` in der Konsole ausgibt.

Erstellen Sie dann eine Variable namens `min` und weisen Sie ihr den Wert von `1` zu, und erstellen Sie eine Variable namens `max`, der Sie den Wert von `100` zuweisen.

Im nächsten Schritt werden Sie eine Zufallszahl zwischen diesen beiden Werten generieren.

# --hints--

Sie sollten einen `console.log()` haben, der die Nachricht `"Now, generate a random number between two values."` ausgibt.

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+number\s+between\s+two\s+values\.\s*\1\);?/);
```

Sie sollten eine Variable namens `min` haben.

```js
assert.isNotNull(min);
```

Ihre `min`-Variable sollte eine Zahl sein.

```js
assert.isNumber(min);
```

Sie sollten die Zahl `1` Ihrer `min`-Variable zuweisen.

```js
assert.equal(min, 1);
```

Sie sollten eine Variable namens `max` haben.

```js
assert.isNotNull(max);
```

Ihre `max`-Variable sollte eine Zahl sein.

```js
assert.isNumber(max);
```

Sie sollten die Zahl `100` Ihrer `max`-Variable zuweisen.

```js
assert.equal(max, 100); 
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);
--fcc-editable-region--

--fcc-editable-region--
```
