---
id: 66ea89c25e1bd871688eba6a
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Bis zu diesem Punkt haben Sie nur mit Gleitkommazahlen und der Zahl `0` gearbeitet.

Im nächsten Teil des Workshops werden Sie überprüfen, wie man mit der `Math.floor()`-Methode arbeitet.

Beginnen Sie damit, eine `console.log()` hinzuzufügen, die die Nachricht `"The Math.floor() method rounds the value down to the nearest whole integer."` protokolliert.

# --hints--

Sie sollten die Nachricht `"The Math.floor() method rounds the value down to the nearest whole integer."` protokollieren.

```js
assert.match(code, /console\.log\((["'])The\s+Math\.floor\(\)\s+method\s+rounds\s+the\s+value\s+down\s+to\s+the\s+nearest\s+whole\s+integer\.\1\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);
--fcc-editable-region--

--fcc-editable-region--
```
