---
id: 66ea7e7399244c6d4b6052a1
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Der nächste Teil des Workshops überprüft, wie die `Math.random()`-Methode funktioniert.

Beginnen Sie damit, eine weitere `console.log()` hinzuzufügen, die den String `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."` in der Konsole ausgibt.

# --hints--

Sie sollten eine `console.log()` haben, die den String `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."` ausgibt.

```js
assert.match(code, /console\.log\((`|'|")\s*The\s+Math\.random\(\)\s+method\s+returns\s+a\s+pseudo\s+random\s+number\s+greater\s+than\s+or\s+equal\s+to\s+0\s+and\s+less\s+than\s+1\.\1\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);
--fcc-editable-region--

--fcc-editable-region--
```
