---
id: 66ea845f54195c7020cc521b
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Die Formel, um eine Zufallszahl zwischen zwei Werten zu erzeugen, lautet wie folgt:

```js
Math.random() * (maximum - minimum) + minimum;
```

Dies ergibt eine Gleitkommazahl zwischen zwei Werten.

Erstellen Sie eine Variable namens `randomNum2` und weisen Sie ihr das Ergebnis der Erzeugung eines Werts zwischen den `min`- und `max`-Werten zu.

Geben Sie dann die Variable `randomNum2` in der Konsole aus, um das Ergebnis zu sehen. Versuchen Sie, die Seite zu aktualisieren, um unterschiedliche Ergebnisse zu erhalten.

# --hints--

Sie sollten eine Variable namens `randomNum2` haben.

```js
assert.isNotNull(randomNum2);
```

Ihre `randomNum2`-Variable sollte eine Zahl sein.

```js
assert.isNumber(randomNum2);
```

Sie sollten das Ergebnis der Erzeugung eines Werts zwischen den `min`- und `max`-Werten zuweisen. Stellen Sie sicher, keinen Wert fest zu codieren. Verwenden Sie die oben angegebene Formel.

```js
assert.isAtLeast(randomNum2, min);
assert.isAtMost(randomNum2, max);

assert.notMatch(code, /randomNum2\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum2\s*=\s*\d+;?/);
```

Sie sollten die Variable `randomNum2` in der Konsole ausgeben.

```js
assert.match(code, /console\.log\(\s*randomNum2\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;
--fcc-editable-region--

--fcc-editable-region--
```
