---
id: 66eb3ec3b07dfe8f898646c0
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Wie Sie in den vorherigen Lektionen gelernt haben, rundet die `Math.floor()`-Methode den Wert auf die nächstkleinere ganze Zahl ab.

```js
const price = 10.99;
Math.floor(price); // 10
```

Erstellen Sie eine Variable namens `numRoundedDown` und weisen Sie ihr das Ergebnis zu, das entsteht, wenn Sie die Gleitkommazahl `6.7` auf die nächstkleinere ganze Zahl abrunden.

Geben Sie dann die `numRoundedDown`-Variable in der Konsole aus, um das Ergebnis zu sehen.

# --hints--

Sie sollten eine `numRoundedDown`-Variable haben.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.numRoundedDown);
```

Die `numRoundedDown`-Variable sollte eine Zahl sein.

```js
assert.isNumber(numRoundedDown);
```

Ihre `numRoundedDown`-Variable sollte dem Ergebnis entsprechen, das entsteht, wenn Sie die Gleitkommazahl `6.7` auf die nächstkleinere ganze Zahl abrunden.

```js
assert.equal(numRoundedDown, 6);
```

Sie sollten den Wert von `6` nicht für die `numRoundedDown`-Variable fest codieren.

```js
const explorer = await __helpers.Explorer(code);
const numRoundedDown = explorer.variables.numRoundedDown;
assert.notEqual(numRoundedDown?.value?.toString(), "6");
```

Sie sollten die `numRoundedDown`-Variable in der Konsole ausgeben.

```js
assert.match(code, /console\.log\(\s*numRoundedDown\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");
--fcc-editable-region--

--fcc-editable-region--
```
