---
id: 66eb488c0129e192542f0865
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie man mit der `Math.ceil()` Methode so arbeitet:

```js
const price = 10.01;
Math.ceil(price); // 11
```

Die `Math.ceil()` Methode rundet den Wert auf die nächsthöhere ganze Zahl auf.

Erstellen Sie eine Variable namens `numRoundedUp` und weisen Sie ihr das Ergebnis zu, das entsteht, wenn Sie die Gleitkommazahl `3.2` auf die nächsthöhere ganze Zahl aufrunden.

Geben Sie dann die `numRoundedUp` Variable in der Konsole aus, um das Ergebnis zu sehen.

# --hints--

Sie sollten eine `numRoundedUp` Variable haben.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.numRoundedUp);
```

Die `numRoundedUp` Variable sollte eine Zahl sein.

```js
assert.isNumber(numRoundedUp);
```

Ihre `numRoundedUp` Variable sollte dem Ergebnis entsprechen, das entsteht, wenn Sie die Gleitkommazahl `3.2` auf die nächstniedrigere ganze Zahl abrunden.

```js
assert.equal(numRoundedUp, 4);
```

Sie sollten den Wert von `4` für die `numRoundedUp` Variable nicht fest codieren.

```js
const explorer = await __helpers.Explorer(code);
const numRoundedUp = explorer.variables.numRoundedUp;
assert.notEqual(numRoundedUp?.value?.toString(), "4");
```

Sie sollten die `numRoundedUp` Variable in der Konsole ausgeben.

```js
assert.match(code, /console\.log\(\s*numRoundedUp\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

--fcc-editable-region--

--fcc-editable-region--
```
