---
id: 66ece147323c91f654bf8dd9
title: Schritt 15
challengeType: 1
dashedName: step-15
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie Sie den maximalen und minimalen Wert aus einem Zahlenbereich wie diesem ermitteln:

```js
Math.max(1, 2, 3, 4, 5); // 5
Math.min(1, 2, 3, 4, 5); // 1
```

Erstellen Sie eine Variable namens `maxNum` und weisen Sie ihr das Ergebnis zu, das die maximale Zahl zwischen den Zahlen `3, 125, 55, 24` findet. Geben Sie dann den Wert von `maxNum` in der Konsole aus.

Erstellen Sie darunter eine Variable namens `minNum` und weisen Sie ihr das Ergebnis zu, das die minimale Zahl zwischen den Zahlen `6, 90, 14, 90, 2` findet. Geben Sie dann den Wert von `minNum` in der Konsole aus.

# --hints--

Sie sollten eine Variable namens `maxNum` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.maxNum);
```

Ihre Variable `maxNum` sollte das Ergebnis der Suche nach der maximalen Zahl zwischen den Zahlen `3, 125, 55, 24` zugewiesen bekommen.

```js
assert.equal(maxNum, 125);
```

Sie sollten den Wert von `125` nicht hartkodiert der Variable `maxNum` zuweisen. Verwenden Sie stattdessen die Methode `Math.max()`.

```js
const explorer = await __helpers.Explorer(code);
assert.notEqual(explorer.variables.maxNum?.value?.toString(), "125");
```

Sie sollten den Wert von `maxNum` in der Konsole ausgeben.

```js
assert.match(code, /console\.log\(\s*maxNum\s*\)/);
```

Sie sollten eine Variable namens `minNum` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.minNum);
```

Ihre Variable `minNum` sollte das Ergebnis der Suche nach der minimalen Zahl zwischen den Zahlen `6, 90, 14, 90, 2` zugewiesen bekommen.

```js
assert.equal(minNum, 2);
```

Sie sollten den Wert von `2` nicht hartkodiert der Variable `minNum` zuweisen. Verwenden Sie stattdessen die Methode `Math.min()`.

```js
const explorer = await __helpers.Explorer(code);
assert.notEqual(explorer.variables.minNum?.value?.toString(), "2");
```

Sie sollten den Wert von `minNum` in der Konsole ausgeben.

```js
assert.match(code, /console\.log\(\s*minNum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

console.log(
  "The Math.round() method rounds the value to the nearest whole integer."
);

const numRounded = Math.round(2.7);
console.log(numRounded);
const numRounded2 = Math.round(11.2);
console.log(numRounded2);

console.log("The Math.max() and Math.min() methods are used to get the maximum and minimum number from a range.");

--fcc-editable-region--

--fcc-editable-region--
```
