---
id: 69218e46f4490b33f2481a1a
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

Die Formel, um eine zufällige ganze Zahl zwischen zwei Werten zu erzeugen, lautet wie folgt:

```js
Math.floor(Math.random() * (maximum - minimum) + minimum);
```

Dies liefert ein Ergebnis, das eine ganze Zahl zwischen zwei Werten ist.

Erstellen Sie eine Variable namens `randomInt` und weisen Sie ihr das Ergebnis der Erzeugung eines Werts zwischen den `min`- und `max`-Werten zu.

Geben Sie dann die `randomInt`-Variable in der Konsole aus, um das Ergebnis zu sehen. Versuchen Sie, die Seite zu aktualisieren, um unterschiedliche Ergebnisse zu erhalten.

# --hints--

Sie sollten eine Variable namens `randomInt` haben.

```js
assert.isNotNull(randomInt);
```

Ihre `randomInt`-Variable sollte eine Zahl sein.

```js
assert.isNumber(randomInt);
```

Sie sollten das Ergebnis der Erzeugung eines Werts zwischen den `min`- und `max`-Werten zuweisen. Stellen Sie sicher, keinen Wert fest zu codieren. Sie sollten die oben angegebene Formel verwenden.

```js
assert.isAtLeast(randomInt, min);
assert.isAtMost(randomInt, max);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomInt\s*=\s*\d+;?/);
```

Sie sollten die `randomInt`-Variable in der Konsole ausgeben.

```js
assert.match(code, /console\.log\(\s*randomInt\s*\);?/);
```


# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");
--fcc-editable-region--

--fcc-editable-region--
```
