---
id: 69218a15fe536ee1be38aefb
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Der nächste Teil des Workshops besteht darin, zu überprüfen, wie man eine zufällige ganze Zahl zwischen zwei Werten erzeugt.

Beginnen Sie damit, einen weiteren `console.log()` hinzuzufügen, der die Nachricht `"Now, generate a random integer between two values."` in der Konsole ausgibt.

Im nächsten Schritt erzeugen Sie eine zufällige ganze Zahl zwischen den Werten der bereits erstellten Variablen `min` und `max`.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten einen `console.log()` haben, der die Nachricht `"Now, generate a random integer between two values."` ausgibt.

```js
assert.equal(getLogs()[7], "Now, generate a random integer between two values.");
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

--fcc-editable-region--

--fcc-editable-region--
```
