---
id: 6852f01116912861c7d1b0e6
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

`__init__` 메서드 안에서, 각 이메일이 자신의 정보를 저장할 수 있도록 매개변수 값을 객체의 속성에 할당해야 합니다.

`__init__` 메서드 안에서 `pass` 키워드를 제거하고 매개변수를 `self.sender`, `self.receiver`, `self.subject`, 그리고 `self.body`에 할당하세요.

# --hints--

`sender` 매개변수를 `self.sender`에 할당해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.sender").is_equivalent("self.sender = sender")`))
    }

})
```

`receiver` 매개변수를 `self.receiver`에 할당해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.receiver").is_equivalent("self.receiver = receiver")`))
    }

})
```

`subject` 매개변수를 `self.subject`에 할당해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.subject").is_equivalent("self.subject = subject")`))
    }

})
```

`body` 매개변수를 `self.body`에 할당해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.body").is_equivalent("self.body = body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        pass
--fcc-editable-region--
```
