---
id: 6852f11d412c29665985a915
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

사용자가 이메일을 보낼 수 있게 되었으니, 받은 이메일을 저장할 방법이 필요합니다. 각 사용자마다 이메일을 모을 받은편지함이 있어야 합니다.

`inbox` 클래스 `User` 메서드에 `__init__` 속성을 추가하고 빈 리스트 `[]`로 초기화하세요. 이 속성은 사용자가 받은 모든 이메일을 저장합니다.

# --hints--

`self.inbox = []` 클래스 `User` 메서드에 `__init__`를 추가해야 합니다.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
```
