---
id: 6852f0c3fcedb964ab959921
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

이제 `read` 속성이 `Email` 클래스에 올바르게 추가되었는지 테스트해 보세요. 이전 단계에서 이미 email 객체가 있으니, `read` 속성을 출력해서 기본적으로 `False`인지 확인하세요.

# --hints--

읽기 속성을 테스트하려면 `email_obj.read`를 출력해야 합니다.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)

--fcc-editable-region--
```
