---
id: 6852f06fdde1b4634a263f03
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

이제 `Email` 클래스를 테스트해 이메일 객체를 만들고 정보가 제대로 저장되는지 확인해 보겠습니다. 몇 가지 속성을 출력해서 모든 것이 잘 작동하는지 검증하세요.

`email_obj`를 발신자로, `alice@example.com`를 수신자로, `bob@example.com`를 제목으로, `Hello`를 본문으로 하여 `Hi Bob!`라는 이름의 이메일 객체를 만드세요. 그런 다음 `sender`와 `subject` 속성을 각각 출력해서 제대로 저장되었는지 확인하세요.

# --hints--

지정된 매개변수로 `email_obj`라는 이름의 이메일 객체를 만들어야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_variable("email_obj").is_equivalent("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`))
    }

})
```

`email_obj.sender`를 출력해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.sender)")`))
    }

})
```

`email_obj.subject`를 출력해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.subject)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--

--fcc-editable-region--
```
