---
id: 6852ef95ccc27660275692cd
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

`Email` 클래스는 이메일이 생성될 때마다 정보를 저장해야 합니다.

`Email` 클래스 안에서 기존 `pass` 키워드를 제거하고 `__init__` 메서드를 만들어 `self`, `sender`, `receiver`, `subject`, `body`를 매개변수로 받으세요.

# --hints--

`__init__` 클래스 안에 `Email` 메서드를 정의해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__init__")`))
    }

})
```

`__init__` 메서드는 `self`, `sender`, `receiver`, `subject`, `body` 매개변수를 가져야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").has_args("self, sender, receiver, subject, body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    pass
--fcc-editable-region--
```
