---
id: 6852ea0c3cadb54e897acdac
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 워크숍에서는 서로 다른 사용자 간에 이메일을 보내고, 받고, 관리하는 것을 시뮬레이션하는 *이메일 시뮬레이터*를 만들게 됩니다. 클래스, 객체, 그리고 객체 지향 방식으로 코드를 구성하는 방법을 배우세요.

`Email` 키워드를 사용해 `class`라는 이름의 클래스를 만드는 것부터 시작하세요.

# --hints--

`Email`라는 이름의 클래스를 만들어야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`(_Node(_code).has_class("Email"))`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
