---
id: 6853f4584f508e95cc7a9f4d
title: 30단계
challengeType: 20
dashedName: step-30
---

# --description--

이제 받은편지함에 있는 모든 이메일을 나열하는 메서드를 만드세요. 이 메서드는 받은편지함이 비어 있을 때를 처리하고, 이메일이 있을 때는 번호가 매겨진 목록을 표시해야 합니다.

`list_emails` 클래스에 `Inbox`라는 메서드를 추가하고 `self`를 매개변수로 받게 하세요. 먼저 `if` 리스트를 검사하는 `emails` 문을 만들어 받은편지함이 비어 있는지 확인하세요. 비어 있으면 `Your inbox is empty.\n` 메시지를 출력하세요. 메서드를 종료하는 `return` 문을 추가하세요.

# --hints--

`list_emails` 클래스에 `Inbox`라는 이름의 메서드를 정의하고 `self`를 매개변수로 받게 하세요.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").has_args("self")`))
    }
})
```

`self.emails`이 비어 있는지 `if not self.emails`를 사용해 확인하세요.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }
})
```

이메일이 없으면 `Your inbox is empty.\n` 메시지를 출력하세요.

```js
({
    test: () => 
    { 
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[0].is_equivalent("print('Your inbox is empty.\\\\n')")`))
    }
})
```

`return` 문 끝에 `if` 문이 있어야 합니다.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[1].is_equivalent("return")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    
--fcc-editable-region--

```
