---
id: 6853f4ded9f84896ff182dd8
title: 31단계
challengeType: 20
dashedName: step-31
---

# --description--

빈 받은편지함 확인 후, 메시지 `\nYour Emails:`를 출력하세요. 그다음, `for` 반복문과 열거를 사용해 받은편지함의 모든 이메일을 순회하세요. 시작 번호가 있는 열거 구문은 다음과 같습니다: `enumerate(iterable, start=start_number)`.

열거를 사용해 이메일에 `1`부터 번호를 매기세요. 반복 변수로 `i` (인덱스)와 `email` (이메일 객체)를 사용하세요.

반복문 안에서, 반복 인덱스 다음에 `.`를 붙이고, 공백과 이메일 객체의 문자열 표현을 다음 형식으로 출력하는 f-문자열을 출력하세요:

```md
i. string_representation
```

여기서 `i`는 인덱스이고 `string_representation`는 이메일 객체의 표현입니다.

# --hints--

헤더 `\nYour Emails:`는 `if` 문장 밖에서 출력해야 합니다.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_body()[1].is_equivalent("print('\\\\nYour Emails:')")`))
    }
})
```

`for` 반복문이 `enumerate(self.emails, start=1)`를 순회해야 합니다.

```js
({
    test: () => 
    {
        runPython(`
_iter = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_iter()
assert _iter.is_equivalent("enumerate(self.emails, start=1)") or _iter.is_equivalent("enumerate(self.emails, 1)")
`)
    }
})
```

`for` 반복문은 `i`와 `email`를 사용해 `enumerate(self.emails, start=1)`를 순회해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_vars().is_equivalent("i, email")`))
    }
})
```

반복문 안에서, 각 이메일 요약을 번호가 붙은 접두사와 `__str__` 메서드에서 포맷된 문자열과 함께 출력해야 합니다. 번호 뒤에 `.`를 추가하는 것도 잊지 마세요.

```js
({
    test: () => 
    {
        assert(runPython(`
        cond = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_body()
        cond.is_equivalent("print(f'{i}. {str(email)}')") or cond.is_equivalent("print(f'{i}. {email}')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        
--fcc-editable-region--
```
