---
id: 6853f90fd03acc9fd75ff860
title: 46단계
challengeType: 20
dashedName: step-46
---

# --description--

사용자는 자신의 User 객체를 통해 받은 편지함을 확인하고, 이메일을 읽고, 이메일을 삭제할 수 있어야 합니다.

받은 편지함을 확인하기 위해 `check_inbox` 클래스에 `User`라는 메서드를 추가하세요. 이 메서드는 사용자의 이름이 포함된 개인화된 헤더를 출력한 다음, 모든 이메일을 표시해야 합니다.

헤더는 `\nName's Inbox:`가 사용자 이름으로 교체된 `Name` 형식이어야 합니다.

헤더를 출력한 후에는 사용자의 받은 편지함에서 `list_emails()` 메서드를 호출하세요.

# --hints--

`check_inbox` 클래스에 `User`를 매개변수로 받는 `self`라는 메서드를 정의해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("check_inbox").has_args("self")`))
    }

})
```

`check_inbox` 내에서 `\n[user]'s Inbox:` 형식으로 사용자의 이름이 포함된 헤더를 출력해야 합니다.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("check_inbox").find_body()[0].is_equivalent('print(f"\\\\n{self.name}\\\'s Inbox:")')`))
    }
})
```

`check_inbox` 내에서 이메일을 표시하기 위해 `self.inbox.list_emails()`를 호출해야 합니다.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("check_inbox").find_body()[1].is_equivalent("self.inbox.list_emails()")`))
    }
})

```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)
        print(f'Email sent from {self.name} to {receiver.name}!\n')

--fcc-editable-region--
    
--fcc-editable-region--

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        del self.emails[actual_index]
        print('Email deleted.\n')

```
