---
id: 6853f2da9b8ed891e9fdc5fc
title: 23단계
challengeType: 20
dashedName: step-23
---

# --description--

이제 발신자와 수신자 정보를 출력해 보겠습니다. 이메일을 누가 보냈고 누가 받았는지 보여 주려면 `display_full_email` 메서드에 다음 형식으로 두 개의 출력문을 추가하세요:

- `From: sender`에서 `sender`가 발신자 이름으로 교체된 것
- `To: receiver`에서 `receiver`가 수신자 이름으로 교체된 것

# --hints--

발신자 이름을 출력할 때 `From: {self.sender.name}`를 사용하여 f-문자열을 사용하세요.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'From: {self.sender.name}')")`));
    }
})
```

수신자 이름을 출력할 때 `To: {self.receiver.name}`를 사용하여 f-문자열을 사용하세요.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'To: {self.receiver.name}')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
