---
id: 68594df7779a164e6a61e8ae
title: 32단계
challengeType: 20
dashedName: step-32
---

# --description--

받은편지함에 특정 이메일을 읽는 메서드가 필요합니다. 사용자가 읽고 싶은 이메일 번호(인덱스 0부터 시작)를 지정하면, 메서드가 해당 이메일의 전체 내용을 보여줍니다.

`Inbox` 클래스에 `read_email`라는 메서드를 추가하세요. 이 메서드는 `index` 매개변수를 받습니다. 먼저 받은편지함이 비어 있는지 확인하고, 비어 있으면 `Inbox is empty.\n` 메시지를 출력하세요. 그 후 메서드를 종료하기 위해 `return` 문을 추가하세요.

# --hints--

`Inbox` 클래스에 `read_email`라는 이름의 메서드를 정의하고 `self`과 `index`를 매개변수로 받도록 하세요.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").has_args("self, index")`));
    }
})
```

`self.emails`가 비어 있는지 `if not self.emails`를 사용해 확인하세요.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`));
    }
})
```

이메일이 없으면 `Inbox is empty.\n`를 출력하고 `return` 하세요.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[0].is_equivalent("print('Inbox is empty.\\\\n')")`));
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[1].is_equivalent("return")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

--fcc-editable-region--
    
--fcc-editable-region--
```
