---
id: 685941422c30cd3dc41b1d0a
title: 22단계
challengeType: 20
dashedName: step-22
---

# --description--

이제 이메일 표시 영역에 헤더를 추가해 보세요. `display_full_email` 메서드에서 `mark_as_read()`을 호출한 후, 명확한 헤더로 이메일 표시를 시작하도록 `\n--- Email ---`를 출력하세요.

# --hints--

`\n--- Email ---` 메서드에서 헤더 `display_full_email`를 출력해야 합니다.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print('\\\\n--- Email ---')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

```
