---
id: 68593ea22eac7e228ed585ce
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

이메일은 읽었는지 여부를 추적해야 합니다. `read` 메서드에 `__init__` 속성을 추가하고, 새 이메일은 기본적으로 읽지 않은 상태이므로 `False`로 설정하세요.

# --hints--

`False` 메서드에서 `self.read`에 `__init__`을 할당해야 합니다.

```js
({ 
    test: () => runPython(`
    assert _Node(_code).find_class("Email").find_function("__init__").find_variable("self.read").is_equivalent("self.read = False")
    `) 
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
```
