---
id: 6859400f8f7f872e50e1238c
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

`User` 클래스에는 다른 사용자에게 이메일을 보내는 메서드가 필요합니다. 이메일을 보낼 때는 새 `Email` 객체를 만들고 수신자의 받은편지함에 추가하세요.

`send_email` 클래스에 `User`라는 메서드를 만드세요. 이 메서드는 `receiver`, `subject`, `body` 매개변수를 받아야 합니다.

메서드 안에서 다음 매개변수 값으로 새 `Email` 객체를 만들고 `email`라는 변수에 할당하세요:

- `sender`: `self` (현재 사용자)
- `receiver`: `receiver` (이메일을 받는 사용자)
- `subject`: `subject` (이메일 제목)
- `body`: `body` (이메일 본문)

# --hints--

`send_email` 클래스에 `User` 메서드를 정의해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").has_function("send_email")`))
    }

})
```

`send_email` 메서드는 `self`, `receiver`, `subject`, `body` 매개변수를 가져야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_args("self, receiver, subject, body")`))
    }

})
```

메서드는 올바른 매개변수로 `Email` 객체를 `email`라는 이름으로 만들어야 합니다.

```js
({
  test: () => {
    assert(runPython(`
cond = _Node(_code).find_class("User").find_function("send_email").find_body()[0]
cond.is_equivalent("email = Email(sender=self, receiver=receiver, subject=subject, body=body)") or cond.is_equivalent("email = Email(self, receiver, subject, body)")
    `));
  }
});
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
```
