---
id: 685c1e32254e5993a5c8823e
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

`Email` 클래스를 처음 설정했으니, 이제 이메일을 보내고 받을 수 있는 사용자를 만들어야 합니다. 각 사용자는 이름을 가지며 이메일 작업을 수행할 수 있습니다.

`User` 클래스를 만들고 `__init__` 메서드를 정의하세요. 이 메서드는 `self`와 `name`를 매개변수로 받습니다. 클래스 내에서 `name` 매개변수를 `self.name`에 할당하세요.

# --hints--

`User` 클래스를 만들어야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_class("User")`))
    }

})
```

`User` 클래스는 `__init__` 메서드를 가져야 하며, 이 메서드는 `self`와 `name`를 매개변수로 받습니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").has_args("self, name")`))
    }

})
```

`name` 매개변수를 `self.name`에 할당해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.name").is_equivalent("self.name = name")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--

--fcc-editable-region--
```
