---
id: 685c1eba966c6396850736ea
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

사용자가 이메일을 보내고 받은편지함을 가질 수 있지만, 받은편지함 작업을 효율적으로 관리하려면 전용 클래스를 만들어야 합니다.

`Inbox` 클래스는 이메일 목록을 저장하고 새 이메일을 추가하고, 모든 이메일을 나열하며, 이를 관리하는 메서드를 제공합니다.

`Inbox`라는 새 클래스를 만들고, `__init__` 메서드에서 `emails`라는 빈 목록을 초기화하세요.

# --hints--

`Inbox` 클래스를 만들어야 합니다.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_class("Inbox")`));
    }
})
```

`Inbox` 클래스는 `__init__`를 매개변수로 받는 `self` 메서드를 가져야 합니다.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").has_args("self")`));
    }
})
```

`self.emails = []`를 `Inbox` 클래스의 `__init__` 메서드에서 초기화해야 합니다.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").find_variable("self.emails").is_equivalent("self.emails = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--

--fcc-editable-region--
```
