---
id: 685cdfd1932d8ac6ad986813
title: 28단계
challengeType: 20
dashedName: step-28
---

# --description--

`__str__` 메서드 내에서 이메일이 읽혔는지 여부를 표시하세요.

조건부 표현식은 한 줄로 조건에 따라 값을 할당할 수 있게 해줍니다.

예시는 다음과 같습니다:

```py
x = 10
y = 'Even' if x % 2 == 0 else 'Odd' # y will be Even
```

메서드 내에서, 이전에 이메일이 읽혔으면 문자열 `Read`을 변수 `status`에, 그렇지 않으면 `Unread`을 할당하도록 조건부 표현식을 사용하세요.

# --hints--

`status` 메서드 내에 `__str__` 변수가 있어야 합니다.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__str__").has_variable("status")`));
    }
})
```

조건부 표현식을 `status` 변수에 할당해야 합니다.

```js
({
    test: () => {
        runPython(`
        import ast
        assert isinstance(_Node(_code).find_class("Email").find_function("__str__").find_variable("status").tree.value, ast.IfExp)`);
    }
})
```

조건부 표현식은 이메일이 읽혔으면 문자열 `Read`을, 그렇지 않으면 `Unread`을 평가해야 합니다.

```js
({
    test: () => {
        runPython(`
        if_exps = [
          "'Read' if self.read else 'Unread'",
          "'Read' if self.read == True else 'Unread'",
          "'Unread' if not self.read else 'Read'",
          "'Unread' if self.read == False else 'Read'"
        ]
        foo = _Node(_code).find_class("Email").find_function("__str__")
        assert any(foo.has_stmt(f"status = {if_exp}") for if_exp in if_exps)
        `);
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

--fcc-editable-region--
    def __str__(self):
        pass
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
