---
id: 685cdabd4cded3b340193f72
title: 24단계
challengeType: 20
dashedName: step-24
---

# --description--

이제 이메일 표시 영역에 제목 줄을 추가하세요. `display_full_email` 메서드에서 이메일 제목을 다음 형식으로 출력하는 print 문을 추가하세요:

- `Subject: subject` 여기서 `subject`는 이메일의 제목으로 교체됩니다

# --hints--

이메일 제목을 출력할 때 `Subject: {self.subject}`를 사용해 f-문자열을 활용하세요.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'Subject: {self.subject}')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
