---
id: 685d2fd1f309b9191f257c55
title: 40단계
challengeType: 20
dashedName: step-40
---

# --description--

이메일 시스템에 타임스탬프를 통합하기 전에 datetime 형식 작업을 연습해 보세요. `datetime.datetime.now()` 함수는 현재 날짜와 시간을 알려 주며, `strftime()` 메서드를 사용해 다양한 방식으로 형식을 지정할 수 있습니다.

`strftime()`가 형식 코드와 함께 작동하는 방식은 다음과 같습니다:

```python
now = datetime.datetime.now()
print(now.strftime("%Y-%m-%d"))  # Output: 2024-03-15 (year-month-day with - separator)
```

`%Y` (년), `%m` (월), `%d` (일) 같은 형식 코드는 `strftime()`에 포함할 내용을 알려 주며, `-` 같은 구분자를 사이에 넣을 수 있습니다.

코드 맨 아래에 `current_time`라는 변수를 만들고 `datetime.datetime.now()`를 할당하세요. 그런 다음 `strftime()`을 사용해 `:`을 구분자로 하여 시:분:초 형식으로 시간을 출력하세요.

다음 형식 코드를 사용하세요: `%H`는 24시간 형식의 시, `%M`는 분, `%S`는 초입니다.

# --hints--

`current_time`라는 변수를 만들고 `datetime.datetime.now()`를 할당해야 합니다.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_stmt("current_time = datetime.datetime.now()")`));
    }
})
```

`strftime()`을 사용해 `%H:%M:%S` 형식으로 시간을 출력해야 합니다.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(current_time.strftime('%H:%M:%S'))")`));
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        del self.emails[actual_index]
        print('Email deleted.\n')

--fcc-editable-region--

--fcc-editable-region--
```
