---
id: 6887550fe36f64e27074475a
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

Alice가 Bob에게 이메일을 보내 이메일 전달이 제대로 작동하는지 확인하세요.

Alice의 `send_email` 메서드를 사용해 제목이 `"Hello"`이고 본문이 `"Hi Bob, how are you?"`인 이메일을 Bob에게 보내세요.

# --hints--

정확한 매개변수로 `alice.send_email()`를 호출해야 합니다.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call('alice.send_email(bob, "Hello", "Hi Bob, how are you?")')`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
alice = User("Alice")
bob = User("Bob")

--fcc-editable-region--
```
