---
id: 688754f3494b17e0b352d67d
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

이제 완성된 이메일 시스템을 테스트할 시간입니다! 사용자 몇 명을 만들고 이메일 기능이 어떻게 작동하는지 확인해 보세요.

이름이 `alice`인 `"Alice"`와 이름이 `bob`인 `"Bob"` 두 개의 User 객체를 만드세요. 이것은 사용자가 어떻게 생성되고 서로 상호작용하는지 보여줍니다.

# --hints--

이름이 `alice`인 `"Alice"`라는 User 객체를 만들어야 합니다.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("alice").is_equivalent('alice = User("Alice")')`));
    }
})
```

이름이 `bob`인 `"Bob"`라는 User 객체를 만들어야 합니다.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("bob").is_equivalent('bob = User("Bob")')`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--

--fcc-editable-region--
```
