---
id: 688754d6d59578df1313cb3c
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

이제 User 클래스와 Inbox 클래스를 연결하여 이메일이 실제로 전달되도록 할 시간입니다!

사용자는 단순한 리스트 대신 적절한 Inbox 객체를 가져야 하며, `send_email` 메서드는 수신자의 인박스로 이메일을 전달해야 합니다.

각 사용자가 이전에 생성된 빈 리스트 대신 실제 Inbox 객체를 받도록 `User` 클래스 `__init__`를 업데이트하세요.

# --hints--

User 클래스를 업데이트하여 `self.inbox = Inbox()`를 초기화해야 합니다.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = Inbox()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []
--fcc-editable-region--

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

class Inbox:
    def __init__(self):
        self.emails = []

```
