---
id: 68879d6de6398893be8c4b7b
title: 25단계
challengeType: 20
dashedName: step-25
---

# --description--

이제 이메일 본문을 추가하여 주요 내용을 완성하세요. `display_full_email` 메서드에서 `Body: body` 형식으로 또 다른 출력문을 추가해 이메일 내용을 보여주세요. 여기서 `body`는 이메일의 내용입니다.

# --hints--

이메일 본문은 `f'Body: {self.body}'`를 사용해 출력하세요.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'Body: {self.body}')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

```
