---
id: 68875554c6c32de788afc64e
title: 21단계
challengeType: 20
dashedName: step-21
---

# --description--

사용자가 이메일을 제대로 읽을 수 있는 방법이 필요합니다.

`display_full_email` 클래스에 `Email` 메서드를 추가하세요. 이 메서드는 `self`만 매개변수로 받습니다. 이 메서드 안에서 누군가 이메일을 볼 때 이메일을 읽음으로 표시하기 위해 `mark_as_read()` 메서드를 호출하세요.

이 메서드는 이메일의 전체 내용을 보기 좋게 표시하는 데 사용됩니다.

# --hints--

`display_full_email` 클래스에 `Email` 메서드를 정의해야 합니다.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").has_function("display_full_email")`));
    }
})
```

`display_full_email` 메서드는 `self`만 매개변수로 받아야 합니다.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_args("self")`));
    }
})
```

이 메서드는 이메일을 읽음으로 표시하기 위해 `self.mark_as_read()`을 호출해야 합니다.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("self.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True
--fcc-editable-region--
    
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
