---
id: 68b936bf4f47d6c29a669f99
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

여전히 `send_email` 메서드 안에서, 수신자의 받은편지함에 이메일을 추가하기 위해 `receive_email`의 `receiver`에 있는 `inbox` 메서드를 호출하세요.

# --hints--

`receiver.inbox.receive_email(email)` 메서드에 `send_email`를 추가해야 합니다.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_stmt("receiver.inbox.receive_email(email)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()
        
--fcc-editable-region--
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)
```
