---
id: 68b9350183ff05bd106c79e4
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

이제 `mark_as_read` 클래스에서 만든 `Email` 메서드를 테스트하세요. 기존 `Email` 객체 `email_obj`에 이 메서드를 사용해 상태를 변경하세요. 이메일을 읽음으로 표시한 후, 변경 사항을 확인하려면 `read`의 `email_obj` 속성을 출력하세요.

# --hints--

이메일을 읽음으로 표시하려면 `mark_as_read`에서 `email_obj` 메서드를 호출해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`))
    }

})
```

`read` 메서드를 호출한 후 `email_obj`의 `mark_as_read` 속성을 출력해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).is_ordered("email_obj.mark_as_read()", "print(email_obj.read)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)

--fcc-editable-region--
```
