---
id: 68b9a011daf902ff1eba40a8
title: 20단계
challengeType: 20
dashedName: step-20
---

# --description--

이메일을 받은 후에는 객체 생성, 메서드 호출, 출력문을 모두 제거하세요.

# --hints--

코드에 `alice = User("Alice")`가 없어야 합니다.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_stmt("alice = User(\\"Alice\\")")`));
    }
})
```

코드에 `bob = User("Bob")`가 없어야 합니다.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_stmt("bob = User(\\"Bob\\")")`));
    }
})
```

코드에 `alice.send_email(bob, "Hello", "Hi Bob, how are you?")`가 없어야 합니다.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("alice.send_email(bob, \\"Hello\\", \\"Hi Bob, how are you?\\")")`));
    }
})
```

코드에 `print(len(bob.inbox.emails))`가 없어야 합니다.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(len(bob.inbox.emails))")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
alice = User("Alice")
bob = User("Bob")

alice.send_email(bob, "Hello", "Hi Bob, how are you?")

print(len(bob.inbox.emails))
--fcc-editable-region--
```
