---
id: 68b930ee0e7b8fb81895e59a
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

이제 이메일을 읽음으로 표시하는 메서드를 추가할 거예요. 누군가 이메일을 열면, 그 이메일의 상태를 읽지 않음에서 읽음으로 바꾸고 싶습니다. 이 메서드는 간단해서 `read` 속성을 `True`으로 설정하기만 하면 됩니다.

`mark_as_read` 클래스에 `Email`라는 메서드를 추가하세요. 이 메서드는 `self`만 매개변수로 받고 `read` 속성을 `True`으로 설정해야 합니다.

# --hints--

`mark_as_read` 클래스에 `Email`라는 메서드를 정의하세요.
이 메서드는 `self`만 매개변수로 받아야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").has_args("self")`))
    }

})
```

메서드는 `self.read`을 `True`으로 설정해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").find_variable("self.read").is_equivalent("self.read = True")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
```
