---
id: 68b9366fae8b13c1278761b1
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

받은편지함이 새 이메일을 받을 수 있는 방법이 필요합니다. 누군가 사용자에게 이메일을 보내면, 그 이메일이 받은편지함에 추가되어야 합니다.

`receive_email` 클래스에 `Inbox`라는 메서드를 추가하세요. 이 메서드는 `self`와 이메일 객체 `email`를 매개변수로 받습니다. 메서드 본문 안에서 `email` 메서드를 사용해 `emails`를 `append` 리스트에 추가하세요.

# --hints--

`receive_email` 클래스에 `Inbox`라는 이름의 메서드를 정의하세요. 이 메서드는 `self`와 `email`를 매개변수로 받습니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_args("self, email")`))
    }
})
```

`email`를 `self.emails`에 추가해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_stmt("self.emails.append(email)")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--
class Inbox:
    def __init__(self):
        self.emails = []

    
--fcc-editable-region--
```
