---
id: 68be6a4447e61c4395ef2359
title: 36단계
challengeType: 20
dashedName: step-36
---

# --description--

받은편지함에도 이메일을 삭제하는 방법이 필요합니다.

Python의 `del` 키워드는 인덱스로 리스트 항목을 삭제하는 데 사용할 수 있습니다.

`delete_email` 클래스에 `Inbox`라는 메서드를 추가하세요. `read_email`처럼 다음을 해야 합니다:

- `index` 매개변수를 받고, 받은편지함이 비어 있는지 확인하며, 비어 있으면 메시지 `Inbox is empty.\n`를 출력하세요. 메서드를 종료하려면 `return`를 사용하세요.

# --hints--

`delete_email` 클래스에 `Inbox`라는 메서드를 정의하고 `self`와 `index`를 매개변수로 받으세요.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").has_args("self,index")`))
    }
})

```

먼저 `if not self.emails:`를 사용해 받은편지함이 비어 있는지 확인해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }
})

```

받은편지함이 비어 있으면 메시지 `Inbox is empty.\n`를 출력하고 메서드를 종료하려면 `return`를 사용하세요.

```js
({
    test: () => 
    {
    assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[0].find_body().is_equivalent("print('Inbox is empty.\\\\n')\\nreturn")`))
}
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()
--fcc-editable-region--

--fcc-editable-region--
```
