---
id: 68dfc5082ea4b4cdcf09dd58
title: 34단계
challengeType: 20
dashedName: step-34
---

# --description--

실제 인덱스를 얻은 후, `actual_index`가 0보다 작거나 `self.emails` 리스트의 길이보다 크거나 같은지 확인하는 if 문을 만드세요. 그렇다면 `Invalid email number.\n` 메시지를 출력하고 `return`를 사용해 메서드를 종료하세요.

# --hints--

`actual_index`가 0보다 작거나 `self.emails` 리스트의 길이보다 크거나 같은지 확인하는 if 문을 만들어야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[1].find_conditions()[0].is_equivalent("actual_index < 0 or actual_index >= len(self.emails)")`))
    }
})

```

인덱스가 유효하지 않으면 `Invalid email number.\n` 메시지를 출력하고 `return`를 사용해 메서드를 종료해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[1].find_body()[0].is_equivalent("print('Invalid email number.\\\\n')")`))
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[1].find_body()[1].is_equivalent("return")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

--fcc-editable-region--
    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1

--fcc-editable-region--

```
