---
id: 63b606f09a14cc1781aea1fb
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

在你的表单中，用户将能够输入一个代表其每日卡路里预算的数字。

创建一个 `label` 元素，为其 `for` 属性设置为 `budget` 以及文本设为 `Budget`，然后创建一个 `input` 元素，并将 `id` 设置为 `budget`。

# --hints--

你应该在 `form` 中创建一个 `label` 元素。

```js
assert.exists(document.querySelector('form label'));
```

你的 `label` 元素应将 `for` 属性设置为 `budget`。

```js
assert.equal(document.querySelector('form label')?.getAttribute('for'), 'budget');
```

你的 `label` 元素应包含文本 `Budget`。

```js
assert.equal(document.querySelector('form label')?.innerText, 'Budget');
```

你应该在 `form` 内创建一个 `input` 元素。

```js
assert.exists(document.querySelector('form input'));
```

你的 `input` 元素应位于 `label` 元素之后。

```js
assert.equal(document.querySelector('form input')?.previousElementSibling?.tagName, "LABEL");
```

你的 `input` 元素应将 `id` 设置为 `budget`。

```js
assert.equal(document.querySelector('form input')?.id, "budget");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
--fcc-editable-region--
        <form id="calorie-counter">

        </form>
--fcc-editable-region--
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
