---
id: 613e275749ebd008e74bb62e
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Наразі `img` має своє значення за замовчуванням, яке занадто велике. У CSS є функція `max`, яка повертає найбільше значення з множини, розділеної комами. Наприклад:

```css
img {
  width: max(250px, 25vw);
}
```

У наведеному прикладі ширина зображення буде 250px, якщо ширина вікна перегляду менша за 1000 пікселів. Якщо ширина вікна перегляду більша за 1000 пікселів, ширина зображення буде 25vw. Це тому, що 25vw дорівнює 25% ширини вікна перегляду.

Масштабуйте зображення, використовуючи його `id` як селектор, і встановіть `width` як максимум між `10rem` або `18vw`.

# --hints--

Ви повинні використати селектор `#logo`, щоб націлитися на елемент `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

Ви повинні задати `img` значення `width` рівне `max(10rem, 18vw)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
