---
id: 614ccc21ea91ef1736b9b578
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Для цього практичного заняття вам надано базовий шаблонний код.

Почніть цю подорож у доступність, додавши атрибут `lang` до вашого елемента `html`. Це допоможе читачам екрана визначити мову сторінки.

Також не забудьте підключити свій файл стилів до сторінки.

# --hints--

Ви повинні надати елементу `html` атрибут `lang`. _Підказка: для англійської мови можна використати значення `en`._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

У вашому елементі `link` має бути елемент `head`.

```js
assert.exists(document.querySelector("head>link"));
```

Ваш елемент `link` повинен мати `rel="stylesheet"` та `href="styles.css"`.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>

  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
