---
id: 655774955b097ea14897db12
title: Passo 28
challengeType: 20
dashedName: step-28
---

# --description--

A distância do nó inicial é zero, porque o algoritmo começa sua avaliação exatamente a partir dali.

Após adicionar `node` a `unvisited` no seu loop, crie uma declaração `if` que seja acionada se o node for igual ao node inicial.Então atribua `0` a esse nó dentro do dicionário `distances`.

# --hints--

Você deve criar uma declaração `if` que seja executada quando `node` for igual a `start`.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+node\s+in\s+graph\s*/);        
    assert(block_body.match(/unvisited\.append\s*\(\s*node\s*\)\s*^\s+if\s+(node\s*==\s*start|start\s*==\s*node)\s*:/m));
  }
})
```

Dentro da sua nova declaração `if` você deve atribuir `0` ao nó no dicionário `distances`.

```js
({ test: () =>  {
   const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+(node\s*==\s*start|start\s*==\s*node)\s*/m);  
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*0/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
