---
id: 6557712d77ce2d9bd7e63afd
title: Passo 18
challengeType: 20
dashedName: step-18
---

# --description--

Adicione outro nó conectado a `B` ao seu grafo e chame-o de `C`.

Modifique seu dicionário existente para representar esse arranjo: adicione outra chave `'C'` a `my_graph` e dê a ela o valor da string `'B'`.

Além disso, altere o valor da chave existente `'B'` para a lista `['A', 'C']` para representar as múltiplas conexões do seu nó `'B'`.

# --hints--

Seu dicionário deve ter 3 chaves — `'A'`, `'B'` e `'C'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C"]
    len(my_graph) == 3 and all(key in my_graph for key in key_list)
  `))
})
```

O valor de `my_graph['A']` deve ser a string `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

`my_graph['B']` deve ser uma lista.

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

O valor de `my_graph['B']` deve ser uma lista contendo `'A'` e `'C'`.

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

O valor de `my_graph['C']` deve ser a string `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["C"] == "B"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': 'A'
}
--fcc-editable-region--
```
