---
id: 65578d0f6c78a0b868a43b9c
title: Passo 40
challengeType: 20
dashedName: step-40
---

# --description--

Dentro do loop `while`, a primeira coisa a fazer é definir o nó atual a ser visitado. Para isso você pode usar a função `min()`. Ele retorna o menor item do iterável passado como argumento.

Remova `pass`, depois crie uma variável chamada `current` e atribua a ela `min(unvisited)`.

# --hints--

Você deve criar uma variável `current` no seu loop `while`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].has_variable("current")`)) })
```

Você deve atribuir `min(unvisited)` à sua variável `current`. Lembre-se de deletar `pass`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].find_bodies()[0].is_equivalent("current = min(unvisited)")`)) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--
    while unvisited:
        pass
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
