---
id: 65577a17564ce8a8e06c1460
title: Passo 35
challengeType: 20
dashedName: step-35
---

# --description--

Compreensões de dicionário também suportam a sintaxe condicional `if`/`else`:

```py
{key: val_1 if condition else val_2 for key in dict}
```

No exemplo acima, `dict` é o dicionário existente. Quando `condition` avalia para `True`, `key` terá o valor `val_1` , caso contrário `val_2`.

Use uma compreensão de dicionário para criar um dicionário baseado em `graph` e atribuí-lo à variável `distances`. Atribua à chave o valor zero se o nó for igual ao nó inicial e infinito caso contrário. Use `float('inf')` para alcançar o último.

# --hints--

Você deve usar a sintaxe de compreensão de dicionário para atribuir um valor à sua variável `distances`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;
    assert(function_body.match(/^\s{4}distances\s*=\s*\{\s*(\w+)\s*:\s*0\s+if\s+(?:\1\s*==\s*start|start\s*==\s*\1)\s+else\s+float\s*\(\s*("|')inf\2\s*\)\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    paths = {node: [] for node in graph}
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
