---
id: 655777060d8ddea6741be1b1
title: Passo 32
challengeType: 20
dashedName: step-32
---

# --description--

Todas as distâncias em `distances` são definidas como infinito, exceto para o nó inicial. A lista `unvisited` contém todos os nós no seu grafo. Mas na verdade, você não precisa daquele loop `for` para alcançar esse resultado.

Remova seu `for` loop com todo o seu corpo.

# --hints--

Você deve remover seu loop `for` e todo o código aninhado.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/(^\s*)distances\s*=\s*\{\s*\}\s*\1print\s*\(\s*f("|')Unvisited:\s*\{\s*unvisited\s*\}\\nDistances:\s\{\s*distances\s*\}\2\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
        else:
            distances[node] = float('inf')
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
