---
id: 655790d113d14dbb727eaf41
title: Passo 45
challengeType: 20
dashedName: step-45
---

# --description--

Uma vez que a distância para um nó é definida dentro do dicionário `distances`, você também precisa acompanhar o caminho até esse nó. Se a distância para o nó na tupla processada foi atualizada, o último item em seu caminho é o próprio nó.

Dentro do seu condicional, aninhe outra declaração `if` que seja acionada quando o último elemento de `paths[node]` for igual a `node`. Use `pass` para preencher o corpo da declaração `if`.

# --hints--

Você deve criar declarações `if` aninhadas que verifiquem se `paths[node][-1]` é igual a `node`. Não esqueça de usar `pass`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)if.*:.*^\1(\s{4})if\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:\s*^\1\2\2pass/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')
```
