---
id: 6557716aadbd2d9c42c0e69a
title: Passo 19
challengeType: 20
dashedName: step-19
---

# --description--

Adicione um último nó, `'D'`, que está conectado com `'A'` e `'C'`.

Modifique seu dicionário para representar essa estrutura. Novamente, use uma lista para representar múltiplas conexões.

# --hints--

Seu dicionário deve ter 4 chaves chamadas `'A'`, `'B'`, `'C'` e `'D'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']` deve ser uma lista.

```js
({ test: () => assert(runPython(`
    type(my_graph["A"]) is list
  `))
})
```

`my_graph['A']` deve ser uma lista contendo `'B'` e `'D'`.

```js
({ test: () => assert(runPython(`
    len(my_graph["A"]) == 2 and "B" in my_graph["A"] and "D" in my_graph["A"]
  `))
})
```

`my_graph['B']` deve ser uma lista.

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

`my_graph['B']` deve ser uma lista contendo `'A'` e `'C'`.

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

`my_graph['C']` deve ser uma lista.

```js
({ test: () => assert(runPython(`
    type(my_graph["C"]) is list
  `))
})
```

`my_graph['C']` deve ser uma lista contendo `'B'` e `'D'`.

```js
({ test: () => assert(runPython(`
    len(my_graph["C"]) == 2 and "B" in my_graph["C"] and "D" in my_graph["C"]
  `))
})
```

`my_graph['D']` deve ser uma lista.

```js
({ test: () => assert(runPython(`
    type(my_graph["D"]) is list
  `))
})
```

`my_graph['D']` deve ser uma lista contendo `'A'` e `'C'`.

```js
({ test: () => assert(runPython(`
    len(my_graph["D"]) == 2 and "A" in my_graph["D"] and "C" in my_graph["D"]
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': ['A', 'C'],
    'C': 'B'
}
--fcc-editable-region--
```
