---
id: 655771d889132f9ccd341060
title: Passo 20
challengeType: 20
dashedName: step-20
---

# --description--

Um grafo é chamado de grafo *weighted* quando suas arestas estão associadas a pesos, representando uma distância, tempo ou outro valor quantitativo.

No seu caso, esses pesos serão as distâncias entre cada nó ou ponto no espaço.Para representar um grafo ponderado você pode modificar seu dicionário, usando uma lista de tuplas para cada valor.

O primeiro elemento na tupla será o nó conectado e o segundo elemento será um número inteiro indicando a distância.

Modifique `my_graph['A']` para uma lista de tuplas, considerando que a distância `A-B` é `3` e a distância `A-D` é `1`.

# --hints--

`my_graph["A"]` deve ser uma lista contendo as tuplas `('B', 3)` e `('D', 1)`.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph` deve ter 4 chaves nomeadas `'A'`, `'B'`, `'C'` e `'D'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': ['B', 'D'],
--fcc-editable-region--    
    'B': ['A', 'C'],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
