---
id: 655791e6cf5e03be3de73451
title: Passo 50
challengeType: 20
dashedName: step-50
---

# --description--

Se você tentar descomentar sua chamada de função, ela não funcionará.Você tem alguns bugs para corrigir.O primeiro acontece porque no `if` aninhado você está tentando acessar um elemento que pode não existir na sua lista `paths[node]`.Então, você precisa ter certeza de que `paths[node]` não está vazio antes de acessar `paths[node][-1]`.

Adicione uma condição adicional ao seu `if` aninhado para garantir que `paths[node]` não esteja vazio antes de acessar `paths[node][-1]`.

# --hints--

Você deve adicionar `paths[node]` como a primeira condição no seu `if` aninhado.Use o operador `and` para combinar suas condições.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/if\s+paths\s*\[\s*node\s*\]\s+and\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--                
                if paths[node][-1] == node:
                    paths[node] = paths[current]
--fcc-editable-region--                   
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
