---
id: 65578cb031cd93b77a285db2
title: Passo 38
challengeType: 20
dashedName: step-38
---

# --description--

Sua função vai explorar todos os nós conectados ao nó inicial. Ele calculará os caminhos mais curtos para todos eles. Então, ele removerá o nó inicial dos nós não visitados. 

Em seguida, o nó vizinho mais próximo será visitado e o processo será repetido até que todos os nós sejam visitados.

A partir de agora, você vai trabalhar no loop principal que explora os nós no grafo. Para evitar problemas com a execução de um loop infinito durante o desenvolvimento do algoritmo, transforme sua chamada de função em um comentário.

# --hints--

Você deve transformar sua chamada de função em um comentário.

```js
({ test: () => assert.match(code, /#\s*shortest_path\s*\(\s*my_graph\s*,\s*("|')A\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
