---
id: 65576ff7888f9e96f52a4be1
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Dicionários armazenam dados na forma de pares *key*-*value*. Uma chave é separada do valor correspondente por dois pontos. E cada par chave-valor é separado do par seguinte por uma vírgula:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}
```

Adicione um novo par chave-valor ao seu dicionário. Use a string `'species'` como a chave e a string `'guinea pig'` como o valor.

# --hints--

Você deve ter uma chave `'species'` com o valor `'guinea pig'` dentro do seu dicionário `copper`.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig"}
  `))
})
```

Seu dicionário `copper` deve ter um único par chave-valor.

```js
({ test: () => assert(runPython(`
    len(copper) == 1
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {}
--fcc-editable-region--
```
