---
id: 6557746aad2844a0cd864e12
title: Passo 27
challengeType: 20
dashedName: step-27
---

# --description--

Enquanto o algoritmo explora o grafo, ele deve manter o controle da menor distância atualmente conhecida entre o nó inicial e os outros nós.

Antes do seu loop `for`, crie uma nova variável chamada `distances` e atribua a ela um dicionário vazio.

# --hints--

Você deve ter uma variável chamada `distances`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=.*(?=^\1for.*:)/ms));
  }
})
```

Sua variável `distances` deve ser um dicionário vazio.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=\s*\{\s*\}.*(?=^\1for.*:)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
