---
id: 65578d4fc3afc3b8f554c882
title: Passo 41
challengeType: 20
dashedName: step-41
---

# --description--

`min()` também aceita um argumento somente por palavra-chave. Passando uma função como um argumento adicional para `min()`, você pode modificar a forma como os itens da lista são comparados.

O resultado da linha que você acabou de escrever na etapa anterior é o nó que vem primeiro em ordem alfabética. Em vez disso, você quer selecionar o nó não visitado que possui a menor distância do nó inicial. 

Passe `key=distances.get` como segundo argumento para sua chamada `min()`. Dessa forma, a comparação ocorrerá dependendo do valor que cada item da lista `unvisited` possui dentro do dicionário `distances`.

# --hints--

Você deve passar `key=distances.get` como o segundo argumento para sua chamada `min()`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s{4})while\s+unvisited\s*:\s*^\1\1current\s*=\s*min\s*\(\s*unvisited\s*,\s*key\s*=\s*distances\.get\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
