---
id: 65577236b056379d5dbc7000
title: Passo 23
challengeType: 20
dashedName: step-23
---

# --description--

Agora você vai começar a desenvolver o algoritmo para calcular o caminho mais curto entre cada nó no seu novo grafo.

Declare uma função vazia chamada `shortest_path`.Use a palavra-chave `pass` para preencher o corpo da função. 

# --hints--

Você deve definir uma função chamada `shortest_path`.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(shortest_path)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}


--fcc-editable-region--
```
