---
id: 655774d01daeeaa1978b99d5
title: Passo 29
challengeType: 20
dashedName: step-29
---

# --description--

No início, todos os outros nós no grafo são considerados a uma distância infinita do nó de origem, porque a distância ainda não foi determinada.

Crie uma cláusula `else` e atribua um valor infinito ao nó no dicionário `distances`.Para isso, use a função `float()` com a string `'inf'` como argumento para gerar um número de ponto flutuante que representa o infinito positivo.

# --hints--

Você deve ter uma cláusula `else`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)if.*:.*^\1else\s*:/ms));
  }
})
```

Você deve atribuir `float('inf')` a `distances[node]` dentro da sua nova cláusula `else`.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, "else");
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*float\s*\(\s*("|')inf\1\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
--fcc-editable-region--
```
