---
id: 65577791ad8c26a7705e2919
title: Passo 34
challengeType: 20
dashedName: step-34
---

# --description--

Com uma compreensão de dicionário, você pode criar um dicionário a partir de um dicionário existente:

```py
{key: val for key in dict}
```

No exemplo acima, `val` é o valor que `key` terá no novo dicionário e `dict` é o dicionário existente.

Você quer acompanhar os caminhos entre o nó inicial e cada outro nó.

Após a variável `distances`, crie uma variável `paths` e atribua a ela um dicionário com todas as chaves de `graph`.Atribua uma lista vazia a cada chave e use uma compreensão de dicionário para construir seu dicionário.

# --hints--

Você deve ter uma variável chamada `paths`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^\s{4}paths\s*\=/m));
  }
})
```

Sua variável `paths` deve usar a sintaxe de compreensão de dicionário para atribuir uma lista vazia a cada nó no grafo.

```js
({ test: () =>  {
    const {function_body} = __helpers.python.getDef(code, "shortest_path");       
    assert(function_body.match(/^\s{4}paths\s*\=\s*\{\s*(\w+)\s*:\s*\[\s*\]\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
