---
id: 6557709b0aee699a6a00528c
title: Passo 16
challengeType: 20
dashedName: step-16
---

# --description--

Por exemplo, um grafo pode ser usado para representar dois pontos no espaço, `A` e `B`, conectados por um caminho.Um grafo como este será composto por dois nós conectados por uma aresta.

Substitua a chave existente `'species'` pelas strings `'A'`.Então, substitua o valor correspondente pela string `'B'` para representar a conexão entre os nós `'A'` e `'B'`.

# --hints--

Seu dicionário deve ter uma chave `'A'`.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

Sua tecla `'A'` deve ter `'B'` como valor.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

Seu dicionário deve ter duas chaves.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
