---
id: 655773b0591c5f9f4045883e
title: Passo 24
challengeType: 20
dashedName: step-24
---

# --description--

O algoritmo começará em um nó especificado. Então ele irá explorar o grafo para encontrar o caminho mais curto entre o nó inicial, ou *source*, e todos os outros nós.

Para isso sua função precisa de dois parâmetros: `graph` e `start`. Adicione-os à sua declaração de função.

# --hints--

Sua função deve receber `graph` e `start` como parâmetros, nesta ordem.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(shortest_path))
    sig == '(graph, start)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path():
    pass
--fcc-editable-region--
```
