---
id: 65578f895f2a65ba7a916804
title: Passo 42
challengeType: 20
dashedName: step-42
---

# --description--

Após a atribuição da variável `current`, crie um loop `for` para iterar sobre as tuplas na lista `graph[current]`. Você precisará de duas variáveis de iteração para isso. Lembre-se de usar `pass` para preencher o corpo do loop.

# --hints--

Você deve criar um loop `for` para iterar sobre os itens tupla na lista `graph[current]`. Use duas variáveis de iteração e não esqueça de adicionar a palavra-chave `pass`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)while\s+unvisited\s*:.*^\1\1for\s+\w+\s*,\s*\w+\s+in\s+graph\s*\[\s*current\s*\]\s*:\s*^\1\1\1pass/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited, key=distances.get)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
