---
id: 65577739f57ecca6c39bb4e9
title: Passo 33
challengeType: 20
dashedName: step-33
---

# --description--

O construtor de tipo `list()` permite que você crie uma lista a partir de um iterável.

Modifique a atribuição da sua variável `unvisited` para usar `list()` e passe `graph` como o iterável.

# --hints--

Você deve usar `list()` para gerar uma lista a partir do dicionário `graph`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/list\s*\(\s*graph\s*\)/));
  }
})
```

Você deve atribuir `list(graph)` a `unvisited`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/unvisited\s*=\s*list\s*\(\s*graph\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
