---
id: 6557924d47c325bf27afbe51
title: Passo 53
challengeType: 20
dashedName: step-53
---

# --description--

O algoritmo está completo mas você pode melhorar a saída. Além disso, você pode fornecer à função um argumento adicional para retornar apenas o caminho entre dois nós.

Adicione `target` como o terceiro parâmetro na declaração da sua função e dê a ele o valor padrão de uma string vazia.

# --hints--

Sua função deve receber três parâmetros: `graph`, `start` e `target`. A ordem importa.

```js
({ test: () => assert.match(code, /^def\s+shortest_path\s*\(\s*graph\s*,\s*start\s*,\s*target\s*=?\s*.*\s*\)\s*:/m) })
```

O parâmetro `target` deve ter o valor padrão de uma string vazia.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(shortest_path))
    sig == "(graph, start, target='')"
  `))
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
--fcc-editable-region--
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current][:]
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')

```
