---
id: 655773f8b8b5db9fc6d0ae76
title: Passo 25
challengeType: 20
dashedName: step-25
---

# --description--

Para acompanhar os nós visitados, você precisa de uma lista de todos os nós no grafo. Uma vez que um nó é visitado, ele será removido dessa lista.

Agora, substitua a palavra-chave `pass` por uma variável chamada `unvisited` e atribua a ela uma lista vazia.

# --hints--

Você deve ter uma variável chamada `unvisited` dentro da função `shortest_path`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").has_variable("unvisited")`)) })
```

Você deve atribuir uma lista vazia à sua variável `unvisited`. Lembre-se de deletar `pass`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_body().is_equivalent("unvisited = []")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    pass
--fcc-editable-region--
```
