---
id: 65579228c669fcbebffd01d5
title: Passo 52
challengeType: 20
dashedName: step-52
---

# --description--

O outro bug é sutil.Quando uma distância menor é encontrada para um nó vizinho, `paths[current]` é atribuído ao caminho do nó vizinho, `paths[node]`. 

Isso significa que ambas as variáveis apontam para a mesma lista.Como listas são mutáveis, quando você adiciona o nó vizinho ao seu caminho, tanto `paths[node]` quanto `paths[current]` são modificados porque são a mesma lista.Isso resulta em caminhos errados, embora as distâncias estejam corretas.

Você pode corrigir esse bug atribuindo uma cópia de `paths[current]` ao caminho do nó vizinho.Para isso você pode usar a sintaxe de slice:

```py
my_list[:]
```

Onde `my_list` é a lista que você quer copiar.Modifique a atribuição existente `paths[node] = paths[current]` dentro do seu bloco `if` cortando `paths[current]`.

# --hints--

Você deve atribuir `paths[current][:]` ao caminho do nó vizinho.

```js
({ test: () => assert.match(code, /^(\s*)if\s+paths\s*\[\s*node\s*\]\s+and\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:\s*^\1\s{4}paths\s*\[\s*node\s*\]\s*=\s*paths\s*\[\s*current\s*\]\s*\[\s*::?\s*\]/ms) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--                
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current]
--fcc-editable-region--                    
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')

```
