---
id: 65661b72d6745ebec6a96923
title: Passo 22
challengeType: 20
dashedName: step-22
---

# --description--

Da mesma forma, modifique as duas listas restantes considerando que a distância `C-D` é `7`.

# --hints--

`my_graph['C']` deve ser uma lista contendo as tuplas `('B', 4)` e `('D', 7)`.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 4), ("D", 7)]
    len(my_graph["C"]) == 2 and all(t in my_graph["C"] for t in tuples) 
  `))
})
```

`my_graph['D']` deve ser uma lista contendo as tuplas `('A', 1)` e `('C', 7)`.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 1), ("C", 7)]
    len(my_graph["D"]) == 2 and all(t in my_graph["D"] for t in tuples)
  `))
})
```

`my_graph` deve ter 4 chaves nomeadas `'A'`, `'B'`, `'C'` e `'D'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']` deve ser uma lista contendo as tuplas `('B', 3)` e `('D', 1)`.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph['B']` deve ser uma lista contendo as tuplas `('A', 3)` e `('C', 4)`.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}
--fcc-editable-region--
```
