---
id: 6578b13757611e2825beb8a5
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

As chaves devem ser únicas dentro de um dicionário e podem ser apenas tipos de dados imutáveis. Isso significa que você não pode usar uma lista ou outro dicionário como chaves.

Adicione outra chave `'age'` ao seu dicionário e dê a ela o número inteiro `2` como valor.

# --hints--

Você deve ter uma chave `'age'` com o valor `2` logo após `'species': 'guinea pig'`, dentro do seu dicionário `copper`. Não esqueça da vírgula.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

Você deve ter dois pares chave-valor no seu dicionário `copper`.

```js
({ test: () => assert(runPython(`
    len(copper) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {'species': 'guinea pig'}
--fcc-editable-region--
```
