---
id: 65789506b30453080f77470c
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Até agora, você já conheceu diferentes tipos de dados:

- Tipos de dados imutáveis, como inteiros, strings, tuplas e Booleans.
- Tipos de dados mutáveis, como listas e dicionários.

Um dicionário é identificado por um par de chaves, `{}`.

Comece criando uma variável chamada `copper` e atribua a ela um dicionário vazio usando um par de chaves, da mesma forma que você criaria uma lista vazia com um par de colchetes.

# --hints--

Você deve ter uma variável chamada `copper`.

```js
({ test: () => assert(__userGlobals.has("copper")) })
```

Sua variável `copper` deve ter o valor de um dicionário vazio.Use um par de chaves para isso.

```js
({ test: () => assert(runPython(`
    copper == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
