---
id: 65797670e0c0d016f17e7660
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Agora, no final do seu código, imprima `copper`.

# --hints--

Você deve imprimir `copper` na parte inferior do seu código.

```js
({ test: () =>
  {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /print\s*\(\s*copper\s*\)\s*$/)
  }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'

--fcc-editable-region--
```
