---
id: 6579cbab9825b8170974c69a
title: Passo 13
challengeType: 20
dashedName: step-13
---

# --description--

Você pode remover um par chave-valor de um dicionário usando a palavra-chave `del`:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

del my_dict['occupation']
```

Logo antes do seu loop `for`, use a palavra-chave `del` para deletar a chave `'age'` e seu valor de `copper`.

# --hints--

Você deve usar a palavra-chave `del` para deletar `copper['age']` antes do loop `for`.

```js
({ test: () => assert.match(code, /^del\s+copper\s*\[\s*("|')age\1\s*\].*^for\s*/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
