---
id: 65796fac81f983127558f3f4
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Agora, modifique sua chamada existente de `print()` para imprimir o valor da chave `age`.


# --hints--

Você deve usar `copper['age']` para acessar o valor da chave `age`.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')age\1\s*\]/) })
```

Você deve chamar `print()` passando `copper['age']` como argumento.

```js
({ test: () => assert.match(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['species'])
--fcc-editable-region--
```
