---
id: 6579717f0920131304286804
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Para adicionar um novo par chave-valor após declarar um dicionário, você pode indicar a chave da mesma forma que acessaria uma chave existente e definir o valor da nova chave usando o operador de atribuição:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['country'] = 'Canada'
```

Exclua sua chamada `print()`. Então, depois da declaração `copper`, adicione a chave `'food'` ao seu dicionário e defina seu valor como `'hay'`.

# --hints--

Você não deve ter `print(copper['age'])` no seu código.

```js
({ test: () => assert.notMatch(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

Você deve adicionar a chave `'food'` a `copper` após declarar o dicionário.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')food\1\s*\]/) })
```

Você deve definir `copper['food']` como `'hay'` após declarar o dicionário.

```js
({ test: () => assert.match(code, /^copper\s*\[\s*("|')food\1\s*\]\s*=\s*("|')hay\2/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['age'])
--fcc-editable-region--
```
