---
id: 6579cd5f6dd62c189e53ddbb
title: Passo 14
challengeType: 20
dashedName: step-14
---

# --description--

Agora que você revisou os aspectos básicos dos dicionários, pode prosseguir para construir o algoritmo do caminho mais curto.

Delete todas as linhas de código após a declaração do dicionário `copper`.

# --hints--

Você deve deletar todas as linhas de código após a declaração do seu dicionário.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse( /copper\s*\[.*?\]|del|for|print\s*\(.*?\)/ms.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'
del copper['age']

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
