---
id: 6579ca0923cfa7162089d2f0
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

A mesma sintaxe pode ser usada para alterar o valor de uma chave existente.

Pouco antes da chamada `print()`, acesse a chave `'species'` e reatribua seu valor para `'Cavia porcellus'`.

# --hints--

Você deve reatribuir o `copper['species']` para a string `'Cavia porcellus'` antes da chamada `print()`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /^copper\s*\[\s*("|')species\1\s*\]\s*=\s*("|')Cavia porcellus\2\s+^print\s*\(\s*copper\s*\)/m)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'

print(copper)
--fcc-editable-region--
```
