---
id: 6579dd49fa8a8e1fd06b85a9
title: Passo 15
challengeType: 20
dashedName: step-15
---

# --description--

Grafos são estruturas de dados que representam relações entre pares de elementos.Estes elementos, chamados *nodes*, podem ser objetos da vida real, entidades, pontos no espaço ou outros.As conexões entre os nós são chamadas de *edges*.

Aqui está uma representação visual de um gráfico:

<img alt="a weighted graph with 6 nodes" src="https://cdn.freecodecamp.org/curriculum/python/graph1-example.png" style="background-color: white; height: 350px; width: auto; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">


Renomeie o dicionário `copper` para `my_graph`.Isto representará o grafo para testar seu algoritmo.

# --hints--

Você deve renomear seu dicionário `copper` para `my_graph`.

```js
({ test: () => assert(__userGlobals.has("my_graph")) })
```

Sua variável `my_graph` deve ser um dicionário.

```js
({ test: () => assert(runPython(`
    type(my_graph) is dict 
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
