---
id: 65affe602a2a828b0382a058
title: Passo 17
challengeType: 20
dashedName: step-17
---

# --description--

Substitua a chave `'age'` pela string `'B'` e defina seu valor como a string `'A'` para representar a conexão entre os nós em ambas as direções.

# --hints--

Seu dicionário deve ter uma chave `'B'`.

```js
({ test: () => assert(runPython(`
    "B" in my_graph
  `))
})
```

A sua tecla `'B'` deve ter `'A'` como valor.

```js
({ test: () => assert(runPython(`
    my_graph["B"] == "A"
  `))
})
```

Seu dicionário deve ter uma chave `'A'`.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

Sua tecla `'A'` deve ter `'B'` como valor.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

Seu dicionário deve ter duas chaves.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'age': 2
}
--fcc-editable-region--
```
