---
id: 65b7d3319ebec69b983fb91b
title: Passo 12
challengeType: 20
dashedName: step-12
---

# --description--

Como você pode ver na saída, `.items()` cria uma estrutura de dados que armazena cada par chave-valor em uma tupla distinta. Para iterar sobre os elementos nesses tuplas você pode adicionar uma segunda variável de loop:

```py
for i, j in dict.items():
    print(i, j)
```

Modifique seu loop `for` para usar duas variáveis de loop e imprimir ambas dentro do corpo do loop.

# --hints--

Você deve modificar seu loop `for` para usar duas variáveis.

```js
({ test: () => assert(runPython(`
import ast
var = _Node(_code).find_for_loops()[0].find_for_vars()
isinstance(var.tree, ast.Tuple)
`)) })
```

Você deve imprimir ambas as variáveis do loop dentro do seu `for` loop.

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print{var}")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.items():
    print(i)
--fcc-editable-region--
```
