---
id: 65b7cf140d34058d7ea8935f
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

Se você quiser iterar sobre os valores das chaves do dicionário, uma forma é usar o método `.values()`.

Modifique seu loop `for` para iterar sobre `copper.values()` em vez de `copper` e observe a saída.

# --hints--

Você deve modificar seu loop `for` para iterar sobre `copper.values()` em vez de `copper`.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.values()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper:
    print(i)
--fcc-editable-region--
```
