---
id: 65b7d25992879698180e6a71
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Finalmente, se você quiser poder percorrer os pares chave-valor, você pode usar o método `.items()`.

Modifique seu loop `for` para iterar sobre `copper.items()` em vez de `copper.values()`.

# --hints--

Você deve modificar seu loop `for` para iterar sobre `copper.items()` em vez de `copper.values()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.items()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.values():
    print(i)
--fcc-editable-region--
```
