---
id: af7588ade1100bde429baf20
title: Construa um Detector de Letra Faltando
challengeType: 26
dashedName: build-a-missing-letter-detector
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter uma função chamada `fearNotLetter`.
1. A função `fearNotLetter` deve aceitar um argumento: uma string que representa um intervalo de letras em ordem alfabética que pode ter uma letra faltando.
1. A função deve encontrar a letra faltando no intervalo de letras passado e retorná-la.
1. Se todas as letras estiverem presentes no intervalo, a função deve retornar `undefined`.

# --hints--

Você deve ter uma função `fearNotLetter`.

```js
assert.isFunction(fearNotLetter);
```

`fearNotLetter("abce")` deve retornar a string `d`.

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

`fearNotLetter("abcdefghjklmno")` deve retornar a string `i`.

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

`fearNotLetter("stvwx")` deve retornar a string `u`.

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

`fearNotLetter("bcdf")` deve retornar a string `e`.

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

`fearNotLetter("abcdefghijklmnopqrstuvwxyz")` deve retornar `undefined`.

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    let letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```
