---
id: 68a9f8378ddba1261a268cce
title: 步骤11
challengeType: 0
dashedName: step-11
---

# --description--

你可以使代码更具可重用性。

创建一个名为 `updateCount` 的函数，该函数以按钮元素作为参数。

此函数应包含更新按钮计数的所有逻辑：查找计数元素、解析当前计数、查看是否达到最大值、递增以及更新显示。

然后更新你的事件监听器回调，调用 `updateCount(happyBtn)`，而不是将逻辑写成内联。

# --hints--

你应该创建一个名为 `updateCount` 的函数。

```js
assert.isFunction(updateCount);
```

`updateCount` 函数应接受一个按钮参数。

```js
assert.equal(updateCount.length, 1);
```

`updateCount` 函数应能正确处理任何按钮元素。

```js
const testBtn = document.createElement("button");
testBtn.innerHTML = '<span class="count">3/10</span>';
updateCount(testBtn);
assert.equal(testBtn.querySelector(".count").textContent, "4/10");
```

你的事件监听器应调用 `updateCount`，并将按钮作为参数。

```js
const spy = __helpers.spyOn(window, "updateCount");
document.querySelector("#happy-btn").click();
assert.equal(spy.calls.length, 1);
assert.equal(spy.calls[0][0], document.querySelector("#happy-btn"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Emoji Reactor</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <main>
    <h1 class="title">How are you feeling today?</h1>
    <p class="description">
      Click on the buttons below to rate your emotions.
    </p>
    <div class="btn-container">
      <button id="happy-btn" class="emoji-btn" aria-label="Happy face emoji">
        <span role="img" aria-hidden="true">😊</span>
        <span class="count">0/10</span>
      </button>
    </div>
  </main>
  <script src="./script.js"></script>
</body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --light-grey: #efefef;
  --white: #ffffff;
  --very-dark-blue: #0a0a23;
  --light-purple: #a78bfa;
  --very-light-purple: #c4b5fd;
  --purple: #8b5cf6;
}

body {
  background-color: var(--very-dark-blue);
  color: var(--light-grey);
  font-family: sans-serif;
}

main {
  text-align: center;
  padding: 10px;
}

.btn-container,
button {
  display: flex;
  justify-content: space-evenly;
  align-items: center;
}

.btn-container {
  flex-direction: column;
}

.emoji-btn {
  width: 70%;
  cursor: pointer;
  color: var(--white);
  background-color: var(--light-purple);
  background-image: linear-gradient(
    to bottom,
    var(--very-light-purple),
    var(--light-purple)
  );
  border: 3px solid var(--purple);
  border-radius: 8px;
  padding: 10px;
  font-size: 1.5rem;
  margin: 10px 0;
  transition: background-color 0.3s ease, transform 0.2s ease;
}

@media (min-width: 768px) {
  .emoji-btn {
    width: 30%;
  }
}

.emoji-btn:hover {
  background-color: var(--purple);
  background-image: none;
}

.title {
  margin-top: 15px;
  font-size: 2rem;
}

.description {
  font-size: 1.4rem;
  margin: 20px 0;
}
```

```js
--fcc-editable-region--
const happyBtn = document.querySelector("#happy-btn");

happyBtn.addEventListener("click", () => {
  const countEl = happyBtn.querySelector(".count");
  let currCount = +countEl.textContent.split("/")[0];
  
  if (currCount === 10) return;
  
  currCount++;
  countEl.textContent = `${currCount}/10`;
})
--fcc-editable-region--
```
