---
id: 68a9fa161498c3324ac06fe1
title: 第 12 步
challengeType: 0
dashedName: step-12
---

# --description--

是时候开始处理其他按钮了。剩余的 `button` 元素已为你添加到超文本标记语言文件中。

对于每个新按钮（`#confused-btn`、`#sad-btn` 和 `#loving-btn`），使用其 ID 通过 `querySelector` 选择它，然后添加一个点击事件监听器，该监听器调用 `updateCount` 并传入该按钮。

使用变量名 `confusedBtn`、`sadBtn`、`lovingBtn`。

# --hints--

你应该创建一个名为 `confusedBtn` 的变量。

```js
assert.exists(confusedBtn);
```

你应该将 `#confused-btn` 元素保存到 `confusedBtn` 变量中。

```js
assert.equal(confusedBtn, document.querySelector('#confused-btn'));
```

你应该创建一个名为 `sadBtn` 的变量。

```js
assert.exists(sadBtn);
```

你应该将 `#sad-btn` 元素保存到 `sadBtn` 变量中。

```js
assert.equal(sadBtn, document.querySelector('#sad-btn'));
```

你应该创建一个名为 `lovingBtn` 的变量。

```js
assert.exists(lovingBtn);
```

你应该将 `#loving-btn` 元素保存到 `lovingBtn` 变量中。

```js
assert.equal(lovingBtn, document.querySelector('#loving-btn'));
```

`#confused-btn` 元素应有一个点击事件监听器，该监听器调用 `updateCount`。

```js
const spy = __helpers.spyOn(window, "updateCount");
document.querySelector("#confused-btn").click();
assert.equal(spy.calls.length, 1);
assert.equal(spy.calls[0][0], document.querySelector("#confused-btn"));
```

`#sad-btn` 元素应有一个点击事件监听器，该监听器调用 `updateCount`。

```js
const spy = __helpers.spyOn(window, "updateCount");
document.querySelector("#sad-btn").click();
assert.equal(spy.calls.length, 1);
assert.equal(spy.calls[0][0], document.querySelector("#sad-btn"));
```

`#loving-btn` 元素应有一个点击事件监听器，该监听器调用 `updateCount`。

```js
const spy = __helpers.spyOn(window, "updateCount");
document.querySelector("#loving-btn").click();
assert.equal(spy.calls.length, 1);
assert.equal(spy.calls[0][0], document.querySelector("#loving-btn"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Emoji Reactor</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <main>
    <h1 class="title">How are you feeling today?</h1>
    <p class="description">
      Click on the buttons below to rate your emotions.
    </p>
    <div class="btn-container">
      <button id="happy-btn" class="emoji-btn" aria-label="Happy face emoji">
        <span role="img" aria-hidden="true">😊</span>
        <span class="count">0/10</span>
      </button>
      <button
        id="confused-btn"
        class="emoji-btn"
        aria-label="Confused face emoji"
      >
        <span role="img" aria-hidden="true">😕</span>
        <span class="count">0/10</span>
      </button>
      <button id="sad-btn" class="emoji-btn" aria-label="Angry face emoji">
        <span role="img" aria-hidden="true">😠</span>
        <span class="count">0/10</span>
      </button>
      <button
        id="loving-btn"
        class="emoji-btn"
        aria-label="Loving face emoji"
      >
        <span role="img" aria-hidden="true">😍</span>
        <span class="count">0/10</span>
      </button>
    </div>
  </main>
  <script src="./script.js"></script>
</body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --light-grey: #efefef;
  --white: #ffffff;
  --very-dark-blue: #0a0a23;
  --light-purple: #a78bfa;
  --very-light-purple: #c4b5fd;
  --purple: #8b5cf6;
}

body {
  background-color: var(--very-dark-blue);
  color: var(--light-grey);
  font-family: sans-serif;
}

main {
  text-align: center;
  padding: 10px;
}

.btn-container,
button {
  display: flex;
  justify-content: space-evenly;
  align-items: center;
}

.btn-container {
  flex-direction: column;
}

.emoji-btn {
  width: 70%;
  cursor: pointer;
  color: var(--white);
  background-color: var(--light-purple);
  background-image: linear-gradient(
    to bottom,
    var(--very-light-purple),
    var(--light-purple)
  );
  border: 3px solid var(--purple);
  border-radius: 8px;
  padding: 10px;
  font-size: 1.5rem;
  margin: 10px 0;
  transition: background-color 0.3s ease, transform 0.2s ease;
}

@media (min-width: 768px) {
  .emoji-btn {
    width: 30%;
  }
}

.emoji-btn:hover {
  background-color: var(--purple);
  background-image: none;
}

.title {
  margin-top: 15px;
  font-size: 2rem;
}

.description {
  font-size: 1.4rem;
  margin: 20px 0;
}
```

```js
function updateCount(btn) {
  const countEl = btn.querySelector(".count");
  let currCount = +countEl.textContent.split("/")[0];
  
  if (currCount === 10) return;
  
  currCount++;
  countEl.textContent = `${currCount}/10`;
}

const happyBtn = document.querySelector("#happy-btn");
--fcc-editable-region--

happyBtn.addEventListener("click", () => updateCount(happyBtn));

--fcc-editable-region--
```
