---
id: 66c051d13a6a20255a963695
title: Compilar uma lista de tarefas estilizada
challengeType: 14
dashedName: build-a-stylized-to-do-list
demoType: onClick
---

# --description--

Neste laboratório, você praticará os diferentes estilos que podem ser aplicados aos links quando passado o mouse por cima, quando ele for focado, clicado e visitado.

Atenda às histórias de usuário abaixo e faça todos os testes passarem para completar o laboratório.

**Histórias de usuário:**

1. Você deve ter uma lista não ordenada com a classe `todo-list`.
2. Dentro da lista não ordenada, você deve ter quatro itens de lista.
3. Dentro de cada item de lista, deve haver:

   - Um elemento `input` com o tipo `checkbox` e `id` definido para um valor único.
   - Um elemento `label` com o atributo `for` definido para o `id` do elemento `input` correspondente.
   - Uma lista não ordenada com a classe `sub-item`.
   - Um item de lista com um elemento de âncora nele. A âncora deve ter a classe `sub-item-link`, um valor `href` válido e um valor `target`, que faça o link abrir em uma nova aba.

4. Você deve remover o estilo de lista padrão do seu `.todo-list`.
5. Os elementos de âncora não devem ter nenhuma decoração de texto padrão.
6. Você deve definir a cor do texto dos links para uma cor de sua escolha.
7. Quando os links são visitados, a cor deve mudar para outra cor de sua escolha.
8. Quando os links são passados o mouse por cima, a cor deve mudar para outra cor de sua escolha.
9.  Quando os links são focados, deve haver um contorno colorido em torno do link.
10. Quando os links são clicados, a cor deve mudar para outra cor de sua escolha.

**Observação:** certifique-se de vincular sua folha de estilo ao HTML e de aplicar o CSS.

# --hints--

Você deve ter uma lista não ordenada com a classe `todo-list`.

```js
assert.exists(document.querySelector('.todo-list'));
```

Você deve remover o estilo de lista padrão do elemento `.todo-list`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.todo-list')?.listStyleType === 'none' || new __helpers.CSSHelp(document).getStyle('.todo-list')?.listStyle === 'none');
```

Você deve ter quatro itens de lista dentro da lista não ordenada.

```js
const li = document.querySelectorAll('.todo-list li');
assert.strictEqual(li.length, 8);
```

Cada item de lista deve conter um elemento `input` com o tipo `checkbox`.

```js
const checkboxes = document.querySelectorAll('.todo-list li input[type="checkbox"]');
assert.strictEqual(checkboxes.length, 4);
```

Todos os elementos `input` devem ter um `id`.

```js
const inputs = document.querySelectorAll('input');
assert(!!inputs.length);
for (let input of inputs) {
    assert.isAbove(input?.id.length, 0);
}
```

Após os elementos `label`, deve haver uma lista não ordenada com a classe `sub-item`.

```js
const labelSbItems = document.querySelectorAll('label+ul.sub-item');
assert.equal(labelSbItems.length, 4);
```

O `li` dentro do `ul` com a classe `sub-item` deve ter um elemento de âncora com a classe `sub-item-link`.

```js
const allItems = document.querySelectorAll('label + ul.sub-item li a.sub-item-link');
assert.equal(allItems.length, 4);
```

Todos os elementos de âncora devem ter um `href` válido.

```js
const links = document.querySelectorAll('a');
assert(!!links.length);
for (let link of links) {
    assert.isAbove(link.href.length, 0);
}
```

Todos os elementos `anchor` devem ter um texto.

```js
const links = document.querySelectorAll('a');
assert(!!links.length);
for (let link of links) {
    assert.isAbove(link.innerText.length, 0);
}
```

Todos os elementos `label` devem ter um atributo `for`.

```js
const labels = document.querySelectorAll('label');
assert(!!labels.length);
for (let label of labels) {
    assert.isAbove(label.htmlFor.length, 0);
}
```

Todos os elementos `label` devem ter algum texto.

```js
const labels = document.querySelectorAll('label');
assert(!!labels.length);
for (let label of labels) {
    assert.isAbove(label.innerText.length, 0);
}
```

Cada elemento `a` deve ter um atributo `target` com o valor de `_blank`.

```js
const anchors = document.querySelectorAll('a');
assert(!!anchors.length);
for (let anchor of anchors) {
    assert.equal(anchor.target, '_blank');
}
```

Os links na página não devem ser sublinhados por padrão.

```js
assert(new __helpers.CSSHelp(document).getStyle('a.sub-item-link')?.textDecoration === 'none' || new __helpers.CSSHelp(document).getStyle('a')?.textDecoration === 'none' ||  new __helpers.CSSHelp(document).getStyle('.sub-item-link')?.textDecoration === 'none' );

```

Os links devem mudar de cor quando passado o mouse por cima.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sub-item-link:hover')?.color || new __helpers.CSSHelp(document).getStyle('a:hover')?.color || new __helpers.CSSHelp(document).getStyle('a.sub-item-link:hover')?.color);
```

Os links devem mudar de cor quando estão sendo clicados.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sub-item-link:active')?.color || new __helpers.CSSHelp(document).getStyle('a:active')?.color || new __helpers.CSSHelp(document).getStyle('a.sub-item-link:active')?.color);
```

Os links devem ter um contorno quando focados.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sub-item-link:focus')?.outline || new __helpers.CSSHelp(document).getStyle('a:focus')?.outline || new __helpers.CSSHelp(document).getStyle('a.sub-item-link:focus')?.outline);

```

Os links devem mudar de cor ao serem visitados.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sub-item-link:visited')?.color || new __helpers.CSSHelp(document).getStyle('a:visited')?.color || new __helpers.CSSHelp(document).getStyle('a.sub-item-link:visited')?.color);

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Styled To-Do List</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Styled To-Do List</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="container">
        <h1>My To-Do List</h1>
        <ul class="todo-list">
            <li>
                <input type="checkbox" id="task1">
                <label for="task1">Explore gaming keyboards</label>
                <ul class="sub-item">
                    <li>
                        <a href="#" target="_blank" class="sub-item-link">Store link</a>
                    </li>
                </ul>
            </li>

            <li>
                <input type="checkbox" id="task2">
                <label for="task2">Finish the report</label>
                <ul class="sub-item">
                    <li>
                        <a href="#" target="_blank" class="sub-item-link">Request Access</a>
                    </li>
                </ul>

            </li>

            <li>
                <input type="checkbox" id="task3">
                <label for="task3">View Phone's Warranty</label>
                <ul class="sub-item">
                    <li>
                        <a href="#" target="_blank" class="sub-item-link">View Receipts</a>

                    </li>
                </ul>
            </li>

            <li>
                <input type="checkbox" id="task4">
                <label for="task4">Check Processor Specs</label>
                <ul class="sub-item">
                    <li>
                        <a href="#" target="_blank" class="sub-item-link">View Model Number</a>

                    </li>
                </ul>
            </li>
        </ul>
    </div>
</body>

</html>
```

```css
body {
    font-family: Arial, sans-serif;
    background-color: #f0f0f0;
    margin: 0;
    padding: 20px;
}

.container {
    max-width: 400px;
    margin: 0 auto;
    background-color: #fff;
    padding: 20px;
    border-radius: 10px;
    box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

h1 {
    text-align: center;
    color: #333;
}

.todo-list {
    list-style-type: none;
    padding: 0;
}

.todo-list li {
    margin-bottom: 15px;
    line-height: 1.6;
    padding: 10px;
    background-color: #e5f1e5;
    border-radius: 10px;
    position: relative;
}

a.sub-item-link {
    margin-left: 10px;
    text-decoration: none;
}

.sub-item-link:link {
    color: blue;
}

.sub-item-link:visited {
    color: purple;
}

.sub-item-link:hover {
    color: red;
}

.sub-item-link:focus {
    outline: 2px solid blue;
    outline-offset: 2px;
}

.sub-item-link:active {
    color: green;
}

.sub-item {
    list-style-type: square;
}
```
