---
id: 66db529d37ad966480ebb633
title: 创建一个午餐选择程序
challengeType: 26
dashedName: build-a-lunch-picker-program
---

# --description--

在本实验室中，你将创建一个管理午餐选项的程序。 你会使用一个数组来存储午餐，向数组中添加或移除项目，并随机选择一个午餐选项。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应创建一个 `lunches` 变量，并为其赋值一个空数组，用于存储午餐项目。

2. 你应创建一个函数 `addLunchToEnd` ，该函数的第一个参数是数组，第二个参数是字符串。 该函数应该：
    - 将字符串添加到数组末尾。
    - 将字串 `[Lunch Item] added to the end of the lunch menu.` 记录到控制台日志，其中 `[Lunch Item]` 是传递给函数的字串。
    - 返回更新后的数组。

3. 你应创建一个函数 `addLunchToStart`，该函数的第一个参数是数组，第二个参数是字符串。 该函数应该：
    - 将字符串添加到数组的开头。
    - 将字串 `[Lunch Item] added to the start of the lunch menu.` 记录到控制台，其中 `[Lunch Item]` 是传递给函数的字串。
    - 返回更新后的数组。

4. 你应创建一个以数组为参数的函数 `removeLastLunch` 。 该函数应该：
   - 从数组中删除最后一个元素。
   - 如果移除成功，向控制台记录字串 `[Lunch Item] removed from the end of the lunch menu.`，其中 `[Lunch Item]` 是从数组末尾移除的元素。
   - 如果数组为空，则在控制台记录字符串 `"No lunches to remove."`。
   - 返回更新后的数组。

5. 你应创建一个以数组为参数的函数 `removeFirstLunch` 。 该函数应该：
   - 从数组中删除第一个元素。
   - 如果移除成功，向控制台记录字串 `[Lunch Item] removed from the start of the lunch menu.`，其中 `[Lunch Item]` 是从数组中移除的元素。
   - 如果数组为空，则在控制台记录字符串 `"No lunches to remove."`。
   - 返回更新后的数组。

6. 你应创建一个以数组为参数的函数 `getRandomLunch` 。 该函数应该：
   - 从数组中随机选择一个元素。
   - 如果成功，向控制台记录字串 `Randomly selected lunch: [Lunch Item]`，其中 `[Lunch Item]` 是数组中的随机元素。
   - 如果数组为空，则在控制台记录字符串 `"No lunches available."`。

7. 你应该创建一个以数组为参数的函数 `showLunchMenu`，且：
   - 如果数组中有元素，则将字串 `Menu items: [Lunch Item], [Lunch Item]...` 记录到控制台，其中每个 `[Lunch item]` 是数组中的一个元素，按顺序排列。
   - 如果数组为空，则在控制台记录字符串 `"The menu is empty."`。

# --hints--

你应声明一个变量 `lunches` 并为其赋值一个空数组来存储午餐项目。

```js
assert.isArray(lunches);
```

你应定义一个函数 `addLunchToEnd`。

```js
assert.isFunction(addLunchToEnd);
```

函数 `addLunchToEnd` 应有两个参数。

```js
assert.lengthOf(addLunchToEnd, 2);
```

`addLunchToEnd(lunches, "Tacos")` 应在控制台记录字符串 `"Tacos added to the end of the lunch menu."`。

```js
const tempArr = [];
const temp = console.log;
const testLunches = []
try {
  console.log = obj => tempArr.push(obj);
  addLunchToEnd(testLunches, "Tacos");
  assert.strictEqual(tempArr[0], "Tacos added to the end of the lunch menu.");
  addLunchToEnd(testLunches, "Pizza");
  assert.strictEqual(tempArr[1], "Pizza added to the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToEnd(["Pizza", "Tacos"], "Burger")` 应返回 `["Pizza", "Tacos", "Burger"]`。

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToEnd(["Pizza", "Tacos"], "Burger"), ["Pizza", "Tacos", "Burger"]);
  assert.deepStrictEqual(addLunchToEnd(["Fries"], "Tacos"), ["Fries", "Tacos"]);
} finally {
  console.log = temp;
}
```

你应定义一个函数 `addLunchToStart` 。

```js
assert.isFunction(addLunchToStart);
```

函数 `addLunchToStart` 应包含两个参数。

```js
assert.lengthOf(addLunchToStart, 2);
```

`addLunchToStart(lunches, "Sushi")` 应在控制台记录字符串 `"Sushi added to the start of the lunch menu."`。

```js
const tempArr = [];
const testLunches = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  addLunchToStart(testLunches, "Sushi");
  assert.strictEqual(tempArr[0], "Sushi added to the start of the lunch menu.");  
  addLunchToStart(testLunches, "Burger");
  assert.strictEqual(tempArr[1], "Burger added to the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToStart(["Burger", "Sushi"], "Pizza")` 应返回 `["Pizza", "Burger", "Sushi"]`。

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToStart(["Burger", "Sushi"], "Pizza"), ["Pizza", "Burger", "Sushi"]);
  assert.deepStrictEqual(addLunchToStart(["Noodles"], "Sushi"), ["Sushi", "Noodles"]);
} finally {
  console.log = temp;
}
```

你应该定义一个函数 `removeLastLunch`。

```js
assert.isFunction(removeLastLunch)
```

函数 `removeLastLunch` 应有一个参数。

```js
assert.lengthOf(removeLastLunch, 1);
```

当输入数组为空时，函数 `removeLastLunch` 应在控制台记录字符串 `"No lunches to remove."`。

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Stew", "Soup", "Toast"])` 应在控制台记录字符串 `"Toast removed from the end of the lunch menu."`。

```js
let testLunches = ["Stew", "Soup", "Toast"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "Toast removed from the end of the lunch menu.");
  testLunches = ["Rice", "Pizza"];
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[1], "Pizza removed from the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Sushi", "Pizza", "Noodles"])` 应返回 `["Sushi", "Pizza"]`。

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeLastLunch(["Sushi", "Pizza", "Noodles"]), ["Sushi", "Pizza"]);
  assert.deepStrictEqual(removeLastLunch(["Pizza", "Burger"]), ["Pizza"]);
} finally {
  console.log = temp;
}
```

你应该定义一个函数 `removeFirstLunch`。

```js
assert.isFunction(removeFirstLunch);
```

函数 `removeFirstLunch` 应该只有一个参数。

```js
assert.lengthOf(removeFirstLunch, 1);
```

当输入数组为空时，函数 `removeFirstLunch` 应在控制台记录字符串 `"No lunches to remove."`。

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Salad", "Eggs", "Cheese"])` 应在控制台记录字符串 `"Salad removed from the start of the lunch menu."`。

```js
let testLunches = ["Salad", "Eggs", "Cheese"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "Salad removed from the start of the lunch menu.");
  testLunches = ["Stew", "Soup", "Toast"]
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[1], "Stew removed from the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Sushi", "Pizza", "Burger"])` 应返回 `["Pizza", "Burger"]`。

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeFirstLunch(["Sushi", "Pizza", "Burger"]), ["Pizza", "Burger"]);
  assert.deepStrictEqual(removeFirstLunch(["Pizza", "Burger"]), ["Burger"]);
} finally {
  console.log = temp;
}
```

`addLunchToEnd`、`addLunchToStart`, `removeLastLunch` 和 `removeFirstLunch` 应在更新后返回作为参数传递的相同数组。

```js
const temp = console.log;
console.log = () => {}
const testLunches = ["Spam"];
try {
  assert.equal(addLunchToEnd(testLunches, "Pizza"), testLunches);
  assert.equal(addLunchToStart(testLunches, "Caviar"), testLunches);
  assert.equal(removeLastLunch(testLunches), testLunches);
  assert.equal(removeFirstLunch(testLunches), testLunches);
} finally {
  console.log = temp;
}
```

你应该定义一个函数 `getRandomLunch`。

```js
assert.isFunction(getRandomLunch);
```

函数 `getRandomLunch` 应该只有一个参数。

```js
assert.lengthOf(getRandomLunch, 1);
```

当输入数组为空时，函数 `getRandomLunch` 应在控制台记录字符串 `"No lunches available."`。

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  getRandomLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches available.");
} finally {
  console.log = temp;
}
```

当输入数组不为空时，函数 `getRandomLunch` 应该在控制台记录格式为 `Randomly selected lunch: [Lunch Item]` 的字串。

```js
const testLunches = ["Fish", "Fries", "Roast"];
const tempRandom = Math.random;
const tempArr = [];
const temp = console.log;

try {
  console.log = obj => tempArr.push(obj);
  // check that it correctly outputs the first item
  Math.random = () => 0;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[0], `Randomly selected lunch: ${testLunches[0]}`);

  // Second item
  Math.random = () => 0.4;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[1], `Randomly selected lunch: ${testLunches[1]}`);

  // Third item
  Math.random = () => 0.8;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[2], `Randomly selected lunch: ${testLunches[2]}`);

} finally {
  Math.random = tempRandom;
  console.log = temp;
}
```

`getRandomLunch` 函数不应更改作为参数传递给它的数组。

```js
const testLunches = ["Fish", "Fries", "Roast"];
let originalLunches = [...testLunches];
const tempRandom = Math.random;

try {
  Math.random = () => 0;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.4;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.8;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  const emptyLunches = [];
  let originalEmpty = [...emptyLunches];
  getRandomLunch(emptyLunches);
  assert.sameOrderedMembers(emptyLunches, originalEmpty);

} finally {
  Math.random = tempRandom;
}
```

你应该定义一个函数 `showLunchMenu`。

```js
assert.isFunction(showLunchMenu);
```

函数 `showLunchMenu` 应该只有一个参数。

```js
assert.lengthOf(showLunchMenu, 1);
```

当输入数组为空时，函数 `showLunchMenu` 应在控制台记录字符串 `"The menu is empty."`。

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], "The menu is empty.");
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Greens", "Corns", "Beans"])` 应在控制台记录字符串 `"Menu items: Greens, Corns, Beans"`。

```js
const testLunches = ["Greens", "Corns", "Beans"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Pizza", "Burger", "Fries", "Salad"])` 应在控制台记录字符串 `"Menu items: Pizza, Burger, Fries, Salad"`。

```js
const testLunches = ["Pizza", "Burger", "Fries", "Salad"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const lunches = [];

// Add a new lunch to the end of the list
function addLunchToEnd(lunches, newLunch) {
  lunches.push(newLunch);
  console.log(`${newLunch} added to the end of the lunch menu.`);
  return lunches;
}

// Add a new lunch to the start of the list
function addLunchToStart(lunches, newLunch) {
  lunches.unshift(newLunch);
  console.log(`${newLunch} added to the start of the lunch menu.`);
  return lunches;
}

// Remove the last lunch from the list
function removeLastLunch(lunches) {
  const removedLunch = lunches.pop();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the end of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Remove the first lunch from the list
function removeFirstLunch(lunches) {
  const removedLunch = lunches.shift();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the start of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Function to get a random lunch from the list
function getRandomLunch(lunches) {
  if (lunches.length === 0) {
    console.log("No lunches available.");
  }
  else {
    const randomIndex = Math.floor(Math.random() * lunches.length);
    const randomLunch = lunches[randomIndex];
    console.log(`Randomly selected lunch: ${randomLunch}`);
  }
}

// Function to display all the lunches in the list
function showLunchMenu(lunches) {
  if (lunches.length === 0) {
    console.log("The menu is empty.");
  }
  else {
    console.log(`Menu items: ${lunches.join(', ')}`);
  }
}

addLunchToEnd(lunches, "Tacos");
addLunchToStart(lunches, "Sushi");
removeLastLunch(lunches);
removeFirstLunch(lunches);
getRandomLunch(lunches);
showLunchMenu(lunches);

```
