---
id: 63b60aaaa65f8922bfce6b7e
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

在你的 `span` 元素中，为 `entry-dropdown` 创建一个 `label` 元素，并为其添加文本 `Add food or exercise:`。 然后创建一个 `select` 元素，将 `id` 设置为 `entry-dropdown`，将 `name` 设置为 `options`。 在其下方，添加一个 `button` 元素，并将 `id` 设置为 `add-entry`，文本为 `Add Entry`。

将 `button` 元素的 `type` 属性设置为 `button`，以防止自动表单提交。

# --hints--

你应该向 `span` 元素添加 `label` 元素。

```js
assert.exists(document.querySelector('.controls > span > label'));
```

你的新 `label` 元素应将 `for` 属性设置为 `entry-dropdown`。

```js
assert.equal(document.querySelector('.controls > span > label')?.getAttribute('for'), 'entry-dropdown');
```

你的新 `label` 元素应包含文本 `Add food or exercise:`。

```js
assert.equal(document.querySelector('.controls > span > label')?.innerText, 'Add food or exercise:');
```

你应该向 `span` 元素添加 `select` 元素。

```js
assert.exists(document.querySelector('.controls > span > select'));
```

你的 `select` 元素应位于 `label` 元素之后。

```js
assert.equal(document.querySelector('.controls > span > select')?.previousElementSibling?.tagName, 'LABEL');
```

你的新 `select` 元素应将 `id` 属性设置为 `entry-dropdown`。

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('id'), 'entry-dropdown');
```

你的新 `select` 元素应将 `name` 属性设置为 `options`。

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('name'), 'options');
```

你的 `select` 元素应为空且具有闭合的 `</select>` 标签。

```js
assert.equal(document.querySelector('.controls > span > select')?.innerHTML?.trim(), '');
```

你应该向 `span` 元素添加一个 `button` 元素。

```js
assert.exists(document.querySelector('.controls > span > button'));
```

你的 `button` 元素应位于 `select` 元素之后。

```js
assert.equal(document.querySelector('.controls > span > button')?.previousElementSibling?.tagName, 'SELECT');
```

你的新 `button` 元素应将 `id` 属性设置为 `add-entry`。

```js
assert.equal(document.querySelector('.controls > span > button')?.getAttribute('id'), 'add-entry');
```

你的新 `button` 元素应该包含文本 `Add Entry`。

```js
assert.equal(document.querySelector('.controls > span > button')?.innerText, 'Add Entry');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--
          <div class="controls">
            <span>

            </span>
          </div>
--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
