---
id: a0b5010f579e69b815e7c5d6
title: 検索と置換
challengeType: 1
forumTopicId: 16045
dashedName: search-and-replace
---

# --description--

指定された引数を使って文の検索と置換を行い、新しい文を返してください。

最初の引数は検索と置換を行う文です。

2番目の引数は置き換える単語（置換前）です。

3番目の引数は2番目の引数を置き換える単語（置換後）です。

**注意:** 置き換える際は、元の単語の最初の文字の大文字・小文字を保持してください。例えば、単語`Book`を`dog`に置き換える場合、`Dog`のように置き換えられるべきです。

# --hints--

`myReplace("Let us go to the store", "store", "mall")`は文字列`Let us go to the mall`を返すべきです。

```js
assert.deepEqual(
  myReplace('Let us go to the store', 'store', 'mall'),
  'Let us go to the mall'
);
```

`myReplace("He is Sleeping on the couch", "Sleeping", "sitting")`は文字列`He is Sitting on the couch`を返すべきです。

```js
assert.deepEqual(
  myReplace('He is Sleeping on the couch', 'Sleeping', 'sitting'),
  'He is Sitting on the couch'
);
```

`myReplace("I think we should look up there", "up", "Down")`は文字列`I think we should look down there`を返すべきです。

```js
assert.deepEqual(
  myReplace('I think we should look up there', 'up', 'Down'),
  'I think we should look down there'
);
```

`myReplace("This has a spellngi error", "spellngi", "spelling")`は文字列`This has a spelling error`を返すべきです。

```js
assert.deepEqual(
  myReplace('This has a spellngi error', 'spellngi', 'spelling'),
  'This has a spelling error'
);
```

`myReplace("His name is Tom", "Tom", "john")`は文字列`His name is John`を返すべきです。

```js
assert.deepEqual(
  myReplace('His name is Tom', 'Tom', 'john'),
  'His name is John'
);
```

`myReplace("Let us get back to more Coding", "Coding", "algorithms")`は文字列`Let us get back to more Algorithms`を返すべきです。

```js
assert.deepEqual(
  myReplace('Let us get back to more Coding', 'Coding', 'algorithms'),
  'Let us get back to more Algorithms'
);
```

# --seed--

## --seed-contents--

```js
function myReplace(str, before, after) {
  return str;
}

myReplace("A quick brown fox jumped over the lazy dog", "jumped", "leaped");
```

# --solutions--

```js
function myReplace(str, before, after) {
  if (before.charAt(0) === before.charAt(0).toUpperCase()) {
    after = after.charAt(0).toUpperCase() + after.substring(1);
  } else {
    after = after.charAt(0).toLowerCase() + after.substring(1);
  }
  return str.replace(before, after);
}
```
