---
id: a105e963526e7de52b219be9
title: ソートされたユニオン
challengeType: 1
forumTopicId: 16077
dashedName: sorted-union
---

# --description--

2つ以上の配列を受け取り、元の配列の順序を保ったまま重複のない新しい配列を返す関数を書いてください。

つまり、すべての配列に存在する値は元の順序で含まれますが、最終的な配列には重複がありません。

ユニークな数値は元の順序で並べられますが、最終的な配列は数値の昇順にはなりません。

例についてはアサーションテストを確認してください。

# --hints--

`uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1])` は `[1, 3, 2, 5, 4]` を返す必要があります。

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

`uniteUnique([1, 2, 3], [5, 2, 1])` は `[1, 2, 3, 5]` を返す必要があります。

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1]), [1, 2, 3, 5]);
```

`uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8])` は `[1, 2, 3, 5, 4, 6, 7, 8]` を返す必要があります。

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8]), [
  1,
  2,
  3,
  5,
  4,
  6,
  7,
  8
]);
```

`uniteUnique([1, 3, 2], [5, 4], [5, 6])` は `[1, 3, 2, 5, 4, 6]` を返す必要があります。

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 4], [5, 6]), [1, 3, 2, 5, 4, 6]);
```

`uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1])` は `[1, 3, 2, 5, 4]` を返す必要があります。

```js
assert.deepEqual(uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

# --seed--

## --seed-contents--

```js
function uniteUnique(arr) {
  return arr;
}

uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]);
```

# --solutions--

```js
function uniteUnique(arr) {
  return [].slice.call(arguments).reduce(function(a, b) {
    return [].concat(
      a, 
      b.filter(function(e, currentIndex) {
        return b.indexOf(e) === currentIndex && a.indexOf(e) === -1;
      }));
  }, []);
}
```
