---
id: a2f1d72d9b908d0bd72bb9f6
title: 人物を作る
challengeType: 1
forumTopicId: 16020
dashedName: make-a-person
---

# --description--

以下のメソッドでオブジェクトコンストラクタを完成させてください。

```js
getFirstName()
getLastName()
getFullName()
setFirstName(first)
setLastName(last)
setFullName(first, last)
```

各メソッドの期待される出力を確認するにはテストを実行してください。これらのメソッドだけがオブジェクトとやり取りする唯一の手段でなければなりません。各テストは新しい`Person`インスタンスを`new Person('Bob', 'Ross')`として宣言します。

# --hints--

関数のシグネチャは変更しないでください。

```js
assert.match(code, /const\s+Person\s*=\s*function\s*\(\s*first\s*,\s*last\s*\)\s*{/);
```

`first`パラメータを再代入しないでください。

```js
assert.notMatch(code, /\bfirst\s*=\s*/);
```

`last`パラメータを再代入しないでください。

```js
assert.notMatch(code, /\blast\s*=\s*/);
```

プロパティを追加しないでください。`Object.keys(Person).length`は常に6を返す必要があります。

```js
const _person = new Person('Bob', 'Ross');
_person.setFirstName('Haskell');
_person.setLastName('Curry');
_person.setFullName('John', 'Smith');
assert.lengthOf(Object.keys(_person), 6);
```

`Person`オブジェクトをインスタンス化できる必要があります。

```js
const _person = new Person('Bob', 'Ross');
assert.instanceOf(_person, Person);
```

`Person`オブジェクトに`firstName`プロパティがあってはいけません。

```js
const _person = new Person('Bob', 'Ross');
assert.notProperty(_person, 'firstName');
```

`Person`オブジェクトに`lastName`プロパティがあってはいけません。

```js
const _person = new Person('Bob', 'Ross');
assert.notProperty(_person, 'lastName');
```

`.getFirstName()`メソッドは文字列`Bob`を返す必要があります。

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getFirstName(), 'Bob');
```

`.getLastName()`は文字列`Ross`を返す必要があります。

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getLastName(), 'Ross');
```

`.getFullName()`メソッドは文字列`Bob Ross`を返す必要があります。

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getFullName(), 'Bob Ross');
```

`.getFullName()`メソッドは`Haskell Ross`を呼び出した後に文字列`.setFirstName('Haskell')`を返す必要があります。

```js
const _person = new Person('Bob', 'Ross');
_person.setFirstName('Haskell');
assert.strictEqual(_person.getFullName(), 'Haskell Ross');
```

`.getFullName()`メソッドは`Bob Curry`を呼び出した後に文字列`.setLastName('Curry')`を返す必要があります。

```js
const _person = new Person('Bob', 'Ross');
_person.setLastName('Curry');
assert.strictEqual(_person.getFullName(), 'Bob Curry');
```

`.getFullName()`メソッドは`Haskell Curry`を呼び出した後に文字列`.setFullName('Haskell', 'Curry')`を返す必要があります。

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getFullName(), 'Haskell Curry');
```

`.getFirstName()`メソッドは`Haskell`を呼び出した後に文字列`.setFullName('Haskell', 'Curry')`を返す必要があります。

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getFirstName(), 'Haskell');
```

`.getLastName()`メソッドは`Curry`を呼び出した後に文字列`.setFullName('Haskell', 'Curry')`を返す必要があります。

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getLastName(), 'Curry');
```

`.getFullName()`メソッドは`Emily Martinez de la Rosa`を呼び出した後に文字列`.setFullName('Emily Martinez', 'de la Rosa')`を返す必要があります。

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getFullName(), 'Emily Martinez de la Rosa');
```

`.getFirstName()`プロパティは`Emily Martinez`を呼び出した後に文字列`.setFullName('Emily Martinez', 'de la Rosa')`を返す必要があります。

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getFirstName(), 'Emily Martinez');
```

`.getLastName()`プロパティは`de la Rosa`を呼び出した後に文字列`.setFullName('Emily Martinez', 'de la Rosa')`を返す必要があります。

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getLastName(), 'de la Rosa');
```

# --seed--

## --seed-contents--

```js
const Person = function(first, last) {
  this.getFullName = function() {
    return "";
  };
  return "";
};
```

# --solutions--

```js
const Person = function(first, last) {
  let firstName = first;
  let lastName = last;

  this.getFirstName = function(){
    return firstName;
  };

  this.getLastName = function(){
    return lastName;
  };

  this.getFullName = function(){
    return firstName + " " + lastName;
  };

  this.setFirstName = function(str){
    firstName = str;
  };


  this.setLastName = function(str){
    lastName = str;
  };

  this.setFullName = function(first, last){
    firstName = first;
    lastName = last;
  };
};
```
