---
id: a3566b1109230028080c9345
title: 範囲内のすべての数字の合計
challengeType: 1
forumTopicId: 16083
dashedName: sum-all-numbers-in-a-range
---

# --description--

2つの数字が入った配列を渡します。その2つの数字と、その間にあるすべての数字の合計を返してください。最小の数字が必ずしも最初に来るとは限りません。

例えば、`sumAll([4,1])` は `10` を返すべきです。なぜなら1から4までのすべての数字の合計（両端を含む）は `10` だからです。

# --hints--

`sumAll([1, 4])` は数値を返すべきです。

```js
assert(typeof sumAll([1, 4]) === 'number');
```

`sumAll([1, 4])` は10を返すべきです。

```js
assert.deepEqual(sumAll([1, 4]), 10);
```

`sumAll([4, 1])` は10を返すべきです。

```js
assert.deepEqual(sumAll([4, 1]), 10);
```

`sumAll([5, 10])` は45を返すべきです。

```js
assert.deepEqual(sumAll([5, 10]), 45);
```

`sumAll([10, 5])` は45を返すべきです。

```js
assert.deepEqual(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--

```js
function sumAll(arr) {
  return 1;
}

sumAll([1, 4]);
```

# --solutions--

```js
function sumAll(arr) {
  var sum = 0;
  arr.sort(function(a,b) {return a-b;});
  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }
  return sum;
}
```
