---
id: a5229172f011153519423690
title: 奇数のフィボナッチ数の合計
challengeType: 1
forumTopicId: 16084
dashedName: sum-all-odd-fibonacci-numbers
---

# --description--

正の整数 `num` が与えられたとき、`num` 以下のすべての奇数のフィボナッチ数の合計を返してください。

フィボナッチ数列の最初の2つの数は0と1です。数列のそれ以降の数は、直前の2つの数の合計です。フィボナッチ数列の最初の7つの数は0、1、1、2、3、5、8です。

たとえば、`sumFibs(10)` は `10` を返すべきです。なぜなら、`10` 以下のすべての奇数のフィボナッチ数は1、1、3、5だからです。

# --hints--

`sumFibs(1)` は数値を返すべきです。

```js
assert(typeof sumFibs(1) === 'number');
```

`sumFibs(1000)` は1785を返すべきです。

```js
assert(sumFibs(1000) === 1785);
```

`sumFibs(4000000)` は4613732を返すべきです。

```js
assert(sumFibs(4000000) === 4613732);
```

`sumFibs(4)` は5を返すべきです。

```js
assert(sumFibs(4) === 5);
```

`sumFibs(75024)` は60696を返すべきです。

```js
assert(sumFibs(75024) === 60696);
```

`sumFibs(75025)` は135721を返すべきです。

```js
assert(sumFibs(75025) === 135721);
```

# --seed--

## --seed-contents--

```js
function sumFibs(num) {
  return num;
}

sumFibs(4);
```

# --solutions--

```js
function sumFibs(num) {
  var a = 1;
  var b = 1;
  var s = 0;
  while (a <= num) {
    if (a % 2 !== 0) {
      s += a;
    }
    a = [b, b=b+a][0];
  }
  return s;
}
```
