---
id: a97fd23d9b809dac9921074f
title: 引数は任意
challengeType: 1
forumTopicId: 14271
dashedName: arguments-optional
---

# --description--

2つの引数を合計する関数を作成してください。もし引数が1つだけ渡された場合は、1つの引数を受け取り合計を返す関数を返してください。

例えば、`addTogether(2, 3)` は `5` を返し、`addTogether(2)` は関数を返すべきです。

この返された関数に1つの引数を渡すと合計が返されます：

```js
var sumTwoAnd = addTogether(2);
```

`sumTwoAnd(3)` は `5` を返します。

どちらかの引数が有効な数値でない場合は、undefined を返してください。

# --hints--

`addTogether(2, 3)` は 5 を返すべきです。

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)` は 53.4 を返すべきです。

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)` は `undefined` を返すべきです。

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)` は `undefined` を返すべきです。

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")` は `undefined` を返すべきです。

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)` は関数を返すべきです。

```js
assert.deepEqual(typeof(addTogether(5)), 'function');
```

`addTogether(5)(7)` は 12 を返すべきです。

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])` は `undefined` を返すべきです。

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")` は `undefined` を返すべきです。

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js
function addTogether() {
  return false;
}

addTogether(2,3);
```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
