---
id: ab306dbdcc907c7ddfc30830
title: ネストされた配列の平坦化
challengeType: 1
forumTopicId: 16079
dashedName: steamroller
---

# --description--

ネストされた配列を平坦化してください。ネストの深さが異なる場合にも対応する必要があります。

# --hints--

`steamrollArray([[["a"]], [["b"]]])` は `["a", "b"]` を返すべきです。

```js
assert.deepEqual(steamrollArray([[['a']], [['b']]]), ['a', 'b']);
```

`steamrollArray([1, [2], [3, [[4]]]])` は `[1, 2, 3, 4]` を返すべきです。

```js
assert.deepEqual(steamrollArray([1, [2], [3, [[4]]]]), [1, 2, 3, 4]);
```

`steamrollArray([1, [], [3, [[4]]]])` は `[1, 3, 4]` を返すべきです。

```js
assert.deepEqual(steamrollArray([1, [], [3, [[4]]]]), [1, 3, 4]);
```

`steamrollArray([1, {}, [3, [[4]]]])` は `[1, {}, 3, 4]` を返すべきです。

```js
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4]);
```

解決策では `Array.prototype.flat()` や `Array.prototype.flatMap()` のメソッドを使わないでください。

```js
assert(!__helpers.removeJSComments(code).match(/\.\s*flat\s*\(/) && !__helpers.removeJSComments(code).match(/\.\s*flatMap\s*\(/));
```

グローバル変数は使わないでください。

```js
steamrollArray([1, {}, [3, [[4]]]])
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4])
```

# --seed--

## --seed-contents--

```js
function steamrollArray(arr) {
  return arr;
}

steamrollArray([1, [2], [3, [[4]]]]);
```

# --solutions--

```js
function steamrollArray(arr) {
  if (!Array.isArray(arr)) {
    return [arr];
  }
  var out = [];
  arr.forEach(function(e) {
    steamrollArray(e).forEach(function(v) {
      out.push(v);
    });
  });
  return out;
}
```
