---
id: af4afb223120f7348cdfc9fd
title: デブリの軌道周期を計算する
challengeType: 1
forumTopicId: 16021
dashedName: map-the-debris
---

# --description--

ケプラーの第三法則によると、円軌道または楕円軌道で互いに回る2つの点質量の軌道周期 $T$ は次の通りです：

$$
T = 2 \pi \sqrt{\frac{a^{3}}{\mu}}
$$

- $a$ は軌道の長半径です
- $μ = GM$ は標準重力パラメータです
- $G$ は重力定数です
- $M$ はより質量の大きい天体の質量です

要素の平均高度を軌道周期（秒単位）に変換した新しい配列を返してください。

配列は `{name: 'name', avgAlt: avgAlt}` 形式のオブジェクトを含みます。

値は最も近い整数に丸めてください。軌道を回る天体は地球です。

地球の半径は6367.4447キロメートルで、地球のGM値は398600.4418 km<sup>3</sup>s<sup>-2</sup>です。

# --hints--

`orbitalPeriod([{name : "sputnik", avgAlt : 35873.5553}])` は `[{name: "sputnik", orbitalPeriod: 86400}]` を返すべきです。

```js
assert.deepEqual(orbitalPeriod([{ name: 'sputnik', avgAlt: 35873.5553 }]), [
  { name: 'sputnik', orbitalPeriod: 86400 }
]);
```

`orbitalPeriod([{name: "iss", avgAlt: 413.6}, {name: "hubble", avgAlt: 556.7}, {name: "moon", avgAlt: 378632.553}])` は `[{name : "iss", orbitalPeriod: 5557}, {name: "hubble", orbitalPeriod: 5734}, {name: "moon", orbitalPeriod: 2377399}]` を返すべきです。

```js
assert.deepEqual(
  orbitalPeriod([
    { name: 'iss', avgAlt: 413.6 },
    { name: 'hubble', avgAlt: 556.7 },
    { name: 'moon', avgAlt: 378632.553 }
  ]),
  [
    { name: 'iss', orbitalPeriod: 5557 },
    { name: 'hubble', orbitalPeriod: 5734 },
    { name: 'moon', orbitalPeriod: 2377399 }
  ]
);
```

# --seed--

## --seed-contents--

```js
function orbitalPeriod(arr) {
  const GM = 398600.4418;
  const earthRadius = 6367.4447;
  return arr;
}

orbitalPeriod([{name : "sputnik", avgAlt : 35873.5553}]);
```

# --solutions--

```js
function orbitalPeriod(arr) {
  const GM = 398600.4418;
  const earthRadius = 6367.4447;
  const TAU = 2 * Math.PI;
  return arr.map(function(obj) {
    return {
      name: obj.name,
      orbitalPeriod: Math.round(TAU * Math.sqrt(Math.pow(obj.avgAlt+earthRadius, 3)/GM))
    };
  });
}
```
