---
id: af7588ade1100bde429baf20
title: 欠けている文字
challengeType: 1
forumTopicId: 16023
dashedName: missing-letters
---

# --description--

渡された文字範囲の中で欠けている文字を見つけて返してください。

範囲内にすべての文字が揃っている場合は `undefined` を返してください。

# --hints--

`fearNotLetter("abce")` は文字列 `d` を返すべきです。

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

`fearNotLetter("abcdefghjklmno")` は文字列 `i` を返すべきです。

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

`fearNotLetter("stvwx")` は文字列 `u` を返すべきです。

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

`fearNotLetter("bcdf")` は文字列 `e` を返すべきです。

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

`fearNotLetter("abcdefghijklmnopqrstuvwxyz")` は `undefined` を返すべきです。

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js
function fearNotLetter(str) {
  return str;
}

fearNotLetter("abce");
```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    var letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```
