---
id: a3566b1109230028080c9345
title: Progetta un algoritmo per sommare tutti i numeri
challengeType: 26
dashedName: lab-sum-all-numbers-algorithm
---

# --description--

In questo laboratorio, dovrai progettare un algoritmo per sommare tutti i numeri.

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione chiamata `sumAll` che accetta un array di due numeri.
1. `sumAll([n, m])` dovrebbe restituire la somma di `n` e `m` più la somma di tutti i numeri compresi tra loro. Il numero più basso non sarà sempre il primo. Per esempio, `sumAll([4,1])` dovrebbe restituire `10` perché la somma di tutti i numeri tra `1` e `4` (inclusi entrambi) è `10`.

# --hints--

Dovresti avere una funzione chiamata `sumAll`.

```js
assert.isFunction(sumAll);
```

`sumAll([1, 4])` dovrebbe restituire un numero.

```js
assert.isNumber(sumAll([1, 4]));
```

`sumAll([1, 4])` dovrebbe restituire `10`.

```js
assert.equal(sumAll([1, 4]), 10);
```

`sumAll([4, 1])` dovrebbe restituire `10`.

```js
assert.equal(sumAll([4, 1]), 10);
```

`sumAll([5, 10])` dovrebbe restituire `45`.

```js
assert.equal(sumAll([5, 10]), 45);
```

`sumAll([10, 5])` dovrebbe restituire `45`.

```js
assert.equal(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function sumAll(arr) {
  let sum = 0;
  
  arr.sort((a, b) => a - b);

  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }

  return sum;
}
```
