---
id: a105e963526e7de52b219be9
title: ユニークでソートされた和集合を実装する
challengeType: 26
dashedName: implement-a-unique-sorted-union
---

# --description--

**目的：** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー：**

1. `uniteUnique` という名前の関数が必要です。
1. `uniteUnique` 関数は2つ以上の配列を引数として受け取る必要があります。
1. 関数は、引数の配列から最初に見つかった順にユニークな値を含む新しい配列を返す必要があります。例えば、`[1, 2, 4], [2, 3, 5]` のような入力に対しては `[1, 2, 4, 3, 5]` のような出力が得られます。

# --hints--

`uniteUnique` 関数が必要です。

```js
assert.isFunction(uniteUnique);
```

`uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1])` は `[1, 3, 2, 5, 4]` を返すべきです。

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

`uniteUnique([1, 2, 3], [5, 2, 1])` は `[1, 2, 3, 5]` を返すべきです。

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1]), [1, 2, 3, 5]);
```

`uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8])` は `[1, 2, 3, 5, 4, 6, 7, 8]` を返すべきです。

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8]), [
  1,
  2,
  3,
  5,
  4,
  6,
  7,
  8
]);
```

`uniteUnique([1, 3, 2], [5, 4], [5, 6])` は `[1, 3, 2, 5, 4, 6]` を返すべきです。

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 4], [5, 6]), [1, 3, 2, 5, 4, 6]);
```

`uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1])` は `[1, 3, 2, 5, 4]` を返すべきです。

```js
assert.deepEqual(uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const uniteUnique = (...arrays) => {
  return arrays.reduce((a, b) => {
    return [
      ...a,
      ...b.filter((e, currentIndex) => b.indexOf(e) === currentIndex && a.indexOf(e) === -1)
    ];
  }, []);
};
```
