---
id: 66f26c32ec6f90df01a44f60
title: Розробіть контактну форму
challengeType: 25
dashedName: design-a-contact-form
demoType: onClick
---

# --description--

**Мета:** виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Ви повинні мати елемент `div` з класом `form-container`, який міститиме форму.

2. Ви повинні мати елемент `form` в межах елемента `.form-container`.

3. Елемент `form` повинен містити наступні елементи та поля введення:

   - елемент `h2` з текстом;
   - поле для введення імені з `type` зі значенням `text`, а також атрибутами `id`, `name` та `required`;
   - поле для введення електронної пошти з `type` зі значенням `email`, а також атрибутами `id`, `name` та `required`;
   - текстову область для введення повідомлення з атрибутами `id`, `name` та `required`;
   - елемент `button` з текстом `"Submit"` та атрибутом `type` зі значенням `submit`.

4. Елементи `input` та `textarea` повинні мати відповідний елемент `label` з атрибутом `for`, значення якого відповідає `id` елемента.

5. Елемент `.form-container` повинен мати колір фону, а також встановлені значення для властивостей `border-radius`, `padding` та `width`.

6. Елементи `label` повинні мати поле та колір шрифту.

7. Елементи `input` та `textarea` повинні мати встановлені значення для:

   - `width`;
   - `padding`;
   - нижнього поля.

8. Елемент `button` повинен мати:

   - колір фону;
   - нестандартний розмір шрифту.

9.  Кнопка повинна мати ефект при наведенні курсора, що змінює колір її фону.

**Примітка:** не забудьте прив’язати таблицю стилів до HTML та застосувати CSS.

# --hints--

Ви повинні мати елемент `div` з класом `form-container`.

```js
const el = document.querySelector('div.form-container');
assert.isNotNull(el);
```

В межах елемента `.form-container` має бути елемент `form`.

```js
const el = document.querySelector('.form-container form');
assert.isNotNull(el);
```

В межах елемента `form` має бути елемент `h2`, який містить текст.

```js
const el = document.querySelector('.form-container form h2');
assert.isNotEmpty(el.innerText);
// assert.isAtLeast(document.querySelector('h1').innerText.length, 1);

```

В межах елемента `form` має бути поле введення з `type` зі значенням `text`, де користувач зможе ввести ім’я.

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert.isNotNull(el);
```

Поле введення `type=text` повинне мати атрибути `id`, `name` та `required`.

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert.exists(el);
assert.isTrue(el.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

В межах елемента `form` має бути поле введення з `type` зі значенням `email`, де користувач зможе ввести електронну пошту.

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert.isNotNull(el);
```

Поле введення `type=email` повинне мати атрибути `id`, `name` та `required`.

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert.isTrue(el?.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

В межах елемента `form` має бути елемент `textarea` з атрибутами `id`, `name` та `required`.

```js
const el = document.querySelector('.form-container form textarea');
assert.isNotNull(el);
assert.isNotNull(document.querySelector('.form-container form textarea[required]'));
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

В межах елемента `form` має бути елемент `button` з атрибутом `type` зі значенням `submit`.

```js
const el = document.querySelector('.form-container form button');
assert.strictEqual(el?.getAttribute('type'), 'submit');
```

Кнопка повинна мати текст `"Submit"`.

```js
const el = document.querySelector('.form-container form button');
assert.strictEqual(document.querySelector('.form-container form button').textContent, 'Submit');
```

Елемент `#name` повинен мати елемент `label` з атрибутом `for`, значення якого відповідає значенню `id` відповідного елемента `input`.

```js
const nameInput = document.querySelector('.form-container form input[type="text"]')?.getAttribute('id');

const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

// check if id has a value
assert.isNotEmpty(nameInput);
let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === nameInput) {
        found = true;
        break;
    }
}

assert.isTrue(found);
```

Елемент `#email` повинен мати елемент `label` з атрибутом `for`, значення якого відповідає значенню `id` відповідного елемента `input`.

```js
const emailInput = document.querySelector('.form-container form input[type="email"]')?.getAttribute('id');
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);
// check if id has a value
assert.isNotEmpty(emailInput);
let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === emailInput) {
        found = true;
        break;
    }
}

assert.isTrue(found);
```

Елемент `#message` повинен мати елемент `label` з атрибутом `for`, значення якого відповідає значенню `id` відповідного елемента `textarea`.

```js
const messageInput = document.querySelector('.form-container form textarea')?.getAttribute('id');
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

// check if id has a value

assert.isNotEmpty(messageInput);

assert.exists(document.querySelector(`label[for="${messageInput}"]`));
```

Мітки не повинні бути порожніми.

```js
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

for (let label of labels) {
    const labelText = label.textContent.trim();
    assert.isNotEmpty(labelText);
}

```

Елемент `.form-container` повинен мати колір фону.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.backgroundColor);
```

Елемент `.form-container` повинен мати властивість `border-radius`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.borderRadius);
```

Елемент `.form-container` повинен мати відступ.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.padding);
```

Елемент `.form-container` повинен мати ширину.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.width);
```

Мітки повинні мати поле та колір шрифту.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('label')?.margin);
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('label')?.color);
```

Елементи `input` повинні мати ширину, відступ та нижнє поле.

```js
// Helper function to extract properties from multiple selectors
function getStyleProperties(selectors, properties) {
  const result = {};
  properties.forEach((property) => {
    result[property] = selectors.map((selector) => selector?.[property]).find(Boolean);
  });
  return result;
}

// Initializing selectors
const selectors = [
  new __helpers.CSSHelp(document).getStyle('input'),
  new __helpers.CSSHelp(document).getStyle('input, textarea'),
  new __helpers.CSSHelp(document).getStyle('textarea, input')
];

// Extract width, padding, and margin properties
const { width } = getStyleProperties(selectors, ['width']);
assert.isNotEmpty(width);

const {
  padding,
  paddingTop,
  paddingBottom,
  paddingLeft,
  paddingRight
} = getStyleProperties(selectors, ['padding', 'paddingTop', 'paddingBottom', 'paddingLeft', 'paddingRight']);

// Assert padding properties
assert.isOk(
  padding ||
  (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

const { marginBottom } = getStyleProperties(selectors, ['marginBottom']);
assert.isNotEmpty(marginBottom);

```

Елемент `textarea` повинен мати ширину, відступ та нижнє поле.

```js
// Helper function to extract properties from multiple selectors
function getStyleProperties(selectors, properties) {
  const result = {};
  properties.forEach((property) => {
    result[property] = selectors.map((selector) => selector?.[property]).find(Boolean);
  });
  return result;
}

// Initializing selectors
const selectors = [
  new __helpers.CSSHelp(document).getStyle('textarea'),
  new __helpers.CSSHelp(document).getStyle('input, textarea'),
  new __helpers.CSSHelp(document).getStyle('textarea, input')
];

// Extract width, padding, and margin properties
const { width } = getStyleProperties(selectors, ['width']);
assert.isNotEmpty(width);

const {
  padding,
  paddingTop,
  paddingBottom,
  paddingLeft,
  paddingRight
} = getStyleProperties(selectors, ['padding', 'paddingTop', 'paddingBottom', 'paddingLeft', 'paddingRight']);

// Assert padding properties
assert.isOk(
  padding ||
  (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

const { marginBottom } = getStyleProperties(selectors, ['marginBottom']);
assert.isNotEmpty(marginBottom);
```

Елемент `button` повинен мати колір фону.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button')?.backgroundColor);
```

Елемент `button` повинен мати розмір шрифту.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button')?.fontSize);
```

Елемент `button` повинен мати ефект при наведенні курсора, що змінює колір фону.

```js
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
const btn_hover = gs('button:hover')?.backgroundColor;
assert.isNotEmpty(btn_hover);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Contact Form</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Styled Form Project</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="form-container">
        <form>
            <h2>Contact Us</h2>
            <label for="name">Name:</label>
            <input type="text" id="name" name="name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" required>

            <label for="message">Message:</label>
            <textarea id="message" name="message" required></textarea>

            <button type="submit">Submit</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background: linear-gradient(45deg, #83a4d4, #b6fbff);
    font-family: Arial, sans-serif;
}

.form-container {
    background-color: white;
    border-radius: 10px;
    padding: 20px;
    width: 300px;
    text-align: center;
}

form h2 {
    margin-bottom: 20px;
    color: #333;
}

label {
    display: block;
    margin: 10px 0 5px;
    color: #555;
}

input{
    width: 100%;
    padding: 10px;
    margin-bottom: 10px;
    border: 1px solid #ddd;
    border-radius: 5px;
    box-sizing: border-box;
    resize: vertical;
}

textarea {
    width: 100%;
    padding: 10px;
    margin-bottom: 10px;
    border: 1px solid #ddd;
    border-radius: 5px;
    box-sizing: border-box;
    resize: vertical;
}

button {
    background-color: #4CAF50;
    color: white;
    padding: 10px 15px;
    border: none;
    border-radius: 5px;
    cursor: pointer;
    font-size: 1em;
}

button:hover {
    background-color: #45a049;
}
```
