---
id: 66faac4139dbbd5dd632c860
title: Crie um formulário de candidatura para vagas de trabalho
challengeType: 25
dashedName: lab-job-application-form
demoType: onClick
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório. 

**Histórias de Usuário:**

1. Você deve ter um elemento `div` com a classe `container`.
1. Dentro do elemento `div`, você deve ter um elemento `form`.
1. O formulário deve conter um elemento `input` com o tipo `text` e o id `name` para inserir o nome completo do usuário.
1. Você deve ter outro elemento `input` com o tipo `email` e o id `email` para inserir o endereço de email do usuário.
1. O formulário deve incluir um elemento `select` com o id `position` que permita aos usuários selecionar uma posição de trabalho.
1. Você deve ter um elemento `fieldset` com a classe `radio-group`.
1. Dentro de `.radio-group` você deve ter um conjunto de elementos `input` com o tipo `radio` e rótulos relevantes para selecionar opções de disponibilidade (por exemplo, Tempo Integral, Meio Período). O `name` do grupo deve ser `availability`.
1. Você deve ter um elemento `textarea` com o id `message` para inserir uma mensagem.
1. Você deve associar cada elemento `input` a um elemento `label`. 
1. Você deve ter um elemento `button` com o tipo `submit` para enviar o formulário.
1. Adicione uma pseudo-classe `:focus` aos elementos `input` e `textarea` para alterar a cor da borda quando estiverem focados.
1. Os elementos `input`, `select` e `textarea` devem ter uma pseudo-classe `:invalid` que altera a cor da borda para vermelho quando uma entrada inválida é detectada.
1. Os elementos `input`, `select` e `textarea` devem ter uma pseudo-classe `:valid` que altera a cor da borda para verde quando uma entrada válida é inserida.
1. O elemento `button` deve ter uma pseudo-classe `:hover` que altera a cor de fundo quando o cursor está sobre ele.
1. Use a pseudo-classe `:checked` em `.radio-group input[type="radio"]` para adicionar uma cor de borda, cor de fundo e uma sombra de caixa quando o botão de rádio estiver selecionado.
1. Use a pseudo-classe `:checked` em botões de rádio para alterar a cor do texto do `label` associado quando a opção estiver selecionada.
1. Adicione uma pseudo-classe `:first-of-type` ao elemento `input` para estilizar o primeiro campo de entrada de forma diferente. (por exemplo, cantos arredondados).

**Nota:** Certifique-se de vincular sua folha de estilos no seu HTML e aplicar seu CSS.

# --hints--

Você deve ter um elemento `div` com a classe `container`.

```js
assert.exists(document.querySelector("div.container"));
```

Dentro do elemento `div`, você deve ter um elemento `form`.

```js
assert.exists(document.querySelector("div.container > form"));
```

O formulário deve conter um elemento `input` com o tipo `text` e o id `name` para inserir o nome completo do usuário.

```js
assert.exists(document.querySelector("div.container > form input#name[type='text']"));
```

Você deve ter outro elemento `input` com o tipo `email` e o id `email` para inserir o endereço de email do usuário.

```js
assert.exists(document.querySelector("div.container > form input#email[type='email']"));
```

O formulário deve incluir um elemento `select` com o id `position` contendo alguns elementos `option`.

```js
assert.exists(document.querySelector("div.container > form select#position"));
assert.isNotEmpty(document.querySelectorAll("div.container > form select#position > option"))
```

Você deve ter um elemento `fieldset` ou `section` com a classe `.radio-group`.

```js
assert.lengthOf(document.querySelectorAll('form fieldset.radio-group, form section.radio-group'), 1)
```

Dentro de `.radio-group` você deve ter um grupo de elementos `input` com o tipo `radio` para selecionar opções de disponibilidade.O grupo `name` deve ser `availability`.

```js
assert.isAtLeast(document.querySelectorAll("div.container > form .radio-group input[type='radio'][name='availability']").length, 2);
assert.isEmpty(document.querySelectorAll("div.container > form .radio-group input[type='radio']:not([name='availability'])"))
```

Você deve ter um elemento `textarea` com o id `message` para inserir uma mensagem.

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form textarea#message"));
```

Você deve associar cada elemento `input` a um elemento `label`. 

```js
const els = document.querySelectorAll('input');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label); 
  assert.equal(label.tagName, "LABEL");
})
```

Você deve ter um elemento `button` com o tipo `submit` para enviar o formulário.

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form button[type='submit']"));
```

Você deve adicionar uma pseudo-classe `:focus` aos elementos `input` e `textarea` para alterar a cor da borda quando estiverem focados. Use um seletor de lista na ordem dada.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input:focus, textarea:focus').getPropertyValue('border-color'))
```

Os elementos `input`, `select` e `textarea` devem ter uma pseudo-classe `:invalid` que altera a cor da borda para vermelho quando uma entrada inválida é detectada. Use um seletor de lista na ordem dada.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:invalid, select:invalid, textarea:invalid').getPropertyValue('border-color'), 'red')
```

Os elementos `input`, `select` e `textarea` devem ter uma pseudo-classe `:valid` que altera a cor da borda para verde quando uma entrada válida é inserida. Use um seletor de lista na ordem dada.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:valid, select:valid, textarea:valid').getPropertyValue('border-color'), 'green')
```

O elemento `button` deve ter uma pseudo-classe `:hover` que altera a cor de fundo quando o cursor está sobre ele.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button:hover').getPropertyValue('background-color'))
```

Você deve usar a pseudo-classe `:checked` em `.radio-group input[type="radio"]` para adicionar uma cor de borda quando o botão de rádio estiver selecionado.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('border-color'));
```

Você deve usar a pseudo-classe `:checked` em `.radio-group input[type="radio"]` para adicionar uma cor de fundo quando o botão de rádio estiver selecionado.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('background-color'));
```

Você deve usar a pseudo-classe `:checked` em `.radio-group input[type="radio"]` para adicionar uma sombra na caixa quando o botão de rádio estiver selecionado.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('box-shadow'));
```

Você deve usar a pseudo-classe `:checked` em botões de rádio para alterar a cor do texto do `label` associado quando a opção estiver selecionada.

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.equal(label.tagName, "LABEL");
  input.checked = false;
  const colorBefore = getComputedStyle(label).color;
  input.checked = true;
  const colorAfter = getComputedStyle(label).color;
  assert.notEqual(colorAfter, colorBefore);
})
```

Adicione uma pseudo-classe `:first-of-type` aos elementos `input` para estilizar o primeiro campo de entrada de forma diferente.

```js
const pseudo = new __helpers.CSSHelp(document).getStyle('input:first-of-type');
assert.exists(pseudo);
assert.isNotEmpty([...pseudo]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application Form</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="container">
        <h1>Job Application Form</h1>
        <form id="job-application">
            <label for="name">Full Name:</label>
            <input type="text" id="name" name="name" placeholder="Enter your name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" placeholder="Enter your email" required>

            <label for="position">Position:</label>
            <select id="position" name="position" required>
                <option value="">Select a position</option>
                <option value="developer">Developer</option>
                <option value="designer">Designer</option>
                <option value="manager">Manager</option>
            </select>

            <fieldset class="radio-group">
              <legend>Availability:</legend>
                <input type="radio" id="fullTime" name="availability" value="fullTime" checked>
                <label for="fullTime">Full-Time</label>
                <input type="radio" id="partTime" name="availability" value="partTime">
                <label for="partTime">Part-Time</label>
            </fieldset>

            <label for="message">Why do you want this job?</label>
            <textarea id="message" name="message" rows="5" placeholder="Write your motivation" required></textarea>

            <button type="submit">Submit Application</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: white;
  padding: 20px;
}

.container {
  max-width: 600px;
  margin: 0 auto;
  background-color: white;
  padding: 20px;
  border-radius: 8px;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
}

label {
  display: block;
  margin-bottom: 5px;
  font-weight: bold;
}

input, select, textarea {
  width: 100%;
  padding: 10px;
  margin-bottom: 15px;
  border: 1px solid #ccc;
  border-radius: 4px;
  box-sizing: border-box;
}

.radio-group {
  display: flex;
  align-items: center;
  margin-top: 15px;
}
.radio-group input[type="radio"] {
  margin-right: 10px;
  -webkit-appearance: none;
  -moz-appearance: none;
  appearance: none;
  width: 20px;
  height: 20px;
  border: 2px solid #ccc;
  border-radius: 50%;
  outline: none;
  transition: border-color 0.2s;
}

.radio-group label {
  margin: 0;
  font-weight: normal;
  cursor: pointer;
}

button {
  width: 100%;
  padding: 10px;
  background-color: green;
  border: none;
  color: white;
  border-radius: 4px;
  font-size: 16px;
  cursor: pointer;
}

input:focus, textarea:focus {
  border-color: blue;
  outline: none;
}

input:invalid, select:invalid, textarea:invalid {
  border-color: red;
}

input:valid, select:valid, textarea:valid {
  border-color: green;
}

.radio-group input[type="radio"]:checked {
  border-color: green;
  background-color: green;
  box-shadow: inset 0 0 0 4px white;
}

button:hover {
  background-color: black;
}

input[type="radio"]:checked + label {
  color: green;
}

button:disabled {
  background-color: grey;
}

input:first-of-type {
  border-top-left-radius: 4px;
  border-top-right-radius: 4px;
}

textarea:nth-child(5) {
  border-bottom-left-radius: 4px;
  border-bottom-right-radius: 4px;
}
```
