---
id: 68c343692285523aef0fc6bf
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

L'attributo `scope` serve a specificare se una cella di intestazione è un'intestazione per una riga, una colonna o un gruppo di righe o colonne. Ecco un esempio:

```html
<th scope="col">Example Header</th>
```

Questo aiuta i lettori di schermo a comprendere la relazione tra le celle di intestazione e quelle di dati.

Per tutti gli elementi `th`, aggiungi un attributo `scope` con valore `col`.

# --hints--

Ogni elemento `th` dovrebbe avere un attributo `scope` con valore `"col"`.

```js
const tableHeaders = document.querySelectorAll("th");
tableHeaders.forEach(header => {
  assert.equal(header?.getAttribute("scope"), "col");
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
      --fcc-editable-region--
        <th>Time</th>
        <th>Track A</th>
        <th>Track B</th>
        <th>Track C</th>
      --fcc-editable-region--
      </tr>
    </thead>
  </table>
</body>
</html>
```
