---
id: 68c34329d430583a6cd5f7f1
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

All'interno del tuo elemento `tr`, aggiungi quattro elementi `th`. Il primo elemento `th` dovrebbe contenere il testo `Time`, il secondo dovrebbe contenere il testo `Track A`, il terzo dovrebbe contenere il testo `Track B` e il quarto dovrebbe contenere il testo `Track C`.

# --hints--

Il tuo elemento `tr` dovrebbe avere quattro elementi `th`.

```js
assert.lengthOf(document.querySelectorAll("tr > th"), 4);
```

Il tuo primo elemento `th` dovrebbe contenere il testo `Time`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(1)")?.textContent.trim(), "Time");
```

Il tuo secondo elemento `th` dovrebbe contenere il testo `Track A`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(2)")?.textContent.trim(), "Track A");
```

Il tuo terzo elemento `th` dovrebbe contenere il testo `Track B`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(3)")?.textContent.trim(), "Track B");
```

Il tuo quarto elemento `th` dovrebbe contenere il testo `Track C`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(4)")?.textContent.trim(), "Track C");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

--fcc-editable-region--
    <thead>
      <tr>
        
      </tr>
    </thead>
--fcc-editable-region--
  </table>
</body>
</html>
```
