---
id: 68c34f4aeb2b1a4170506751
title: Passaggio 13
challengeType: 0
dashedName: step-13
---

# --description--

Successivamente, aggiungi una quinta riga alla tabella. Inizia aggiungendo un altro elemento `tr`. All’interno di quell’elemento `tr`, aggiungi un elemento `th` con un attributo `scope` impostato su `"row"` e il contenuto testuale `12:30 PM`. Poi, sotto quell’elemento `th`, aggiungi un elemento `td` con il contenuto testuale `Lunch Break`. Il tuo elemento `td` dovrebbe avere anche un attributo `colspan` impostato su `3` in modo che si estenda su tutte e tre le tracce.

# --hints--

Dovresti avere un quinto elemento `tr` all’interno dell’elemento `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 5);
```

Dovresti avere un elemento `th` all’interno del quinto elemento `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("th"));
```

Il tuo elemento `th` dovrebbe avere il contenuto testuale `12:30 PM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.textContent.trim(), "12:30 PM");
```

Il tuo elemento `th` dovrebbe avere un attributo `scope` impostato su `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.getAttribute("scope"), "row");
```

Dovresti avere un elemento `td` all’interno del quinto elemento `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("td"));
```

Il tuo elemento `td` dovrebbe avere il contenuto testuale `Lunch Break`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.textContent.trim(), "Lunch Break");
``` 

Il tuo elemento `td` dovrebbe avere un attributo `colspan` impostato su `3`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
  </table>
</body>
</html>
```
