---
id: 68c3482278a14e3cf225e897
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

All'interno del tuo elemento `tr`, aggiungi un elemento `th` con il testo `9:00 AM`. Poi, sotto quell'elemento `th`, aggiungi tre elementi `td` con il testo `Keynote: Tech Future`, `Intro to Web Dev` e `UX for All`. 

# --hints--

Dovresti avere un elemento `th` all'interno del tuo elemento `tr`.

```js
assert.exists(document.querySelector("tbody tr th"));
```

Il tuo elemento `th` dovrebbe avere come contenuto testuale `9:00 AM`.

```js
assert.equal(document.querySelector("tbody tr th")?.textContent.trim(), "9:00 AM");
```

Dovresti avere un totale di tre elementi `td` sotto il tuo elemento `th`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.lengthOf(tdElements, 3);
```

Il tuo primo elemento `td` dovrebbe avere come contenuto testuale `Keynote: Tech Future`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[0]?.textContent.trim(), "Keynote: Tech Future");
```

Il tuo secondo elemento `td` dovrebbe avere come contenuto testuale `Intro to Web Dev`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[1]?.textContent.trim(), "Intro to Web Dev");
``` 

Il tuo terzo elemento `td` dovrebbe avere come contenuto testuale `UX for All`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[2]?.textContent.trim(), "UX for All");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
