---
id: 68c34e201b347c40acda8111
title: Passaggio 12
challengeType: 0
dashedName: step-12
---

# --description--

Ora è il momento di aggiungere un'altra riga alla tabella.

Inizia aggiungendo un altro elemento `tr`. All'interno di quell'elemento `tr`, aggiungi un elemento `th` con un attributo `scope` impostato su `"row"` e il contenuto testuale `11:30 AM`.

Poi, aggiungi tre elementi `td` con il seguente contenuto testuale:

- `AR/VR in Education`
- `JavaScript Fundamentals`
- `Design Systems at Scale`

# --hints--

Dovresti avere un quarto elemento `tr` all'interno dell'elemento `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 4);
```

Dovresti avere un elemento `th` all'interno del quarto elemento `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[3].querySelector("th"));
```

Il tuo elemento `th` dovrebbe avere il contenuto testuale `11:30 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.textContent.trim(), "11:30 AM");
```

Il tuo elemento `th` dovrebbe avere un attributo `scope` impostato su `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.getAttribute("scope"), "row");
```

Il tuo quarto elemento `tr` dovrebbe avere tre elementi `td`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[3].querySelectorAll("td"), 3);
```

Il tuo primo elemento `td` dovrebbe avere il contenuto testuale `AR/VR in Education`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[0]?.textContent.trim(), "AR/VR in Education");
```

Il tuo secondo elemento `td` dovrebbe avere il contenuto testuale `JavaScript Fundamentals`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[1]?.textContent.trim(), "JavaScript Fundamentals");
```   

Il tuo terzo elemento `td` dovrebbe avere il contenuto testuale `Design Systems at Scale`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[2]?.textContent.trim(), "Design Systems at Scale");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
