---
id: 68c3420cb0176439a95af1b8
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Il passaggio successivo è aggiungere l'intestazione della tabella.

Inizia aggiungendo un elemento `thead` sotto l'elemento `caption`. All'interno del `thead`, aggiungi un elemento `tr`.

# --hints--

Dovresti avere un elemento `thead`.

```js
assert.exists(document.querySelector("table thead"));
```

Dovresti avere un elemento `tr` all'interno del `thead`.

```js
assert.exists(document.querySelector("table thead tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    --fcc-editable-region--
        
    --fcc-editable-region--
  </table>
</body>
</html>
```
