---
id: 68c346d03709343bff132707
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

Nei prossimi passaggi, costruirai il corpo della tabella.

Inizia aggiungendo un elemento `tbody` sotto l'elemento `thead`. Poi, all'interno di `tbody`, aggiungi un elemento `tr`.

# --hints--

Dovresti avere un elemento `tbody` sotto l'elemento `thead`.

```js
assert.exists(document.querySelector("tbody"));
```

Dovresti avere un elemento `tr` all'interno dell'elemento `tbody`.

```js
assert.exists(document.querySelector("tbody tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>
    
    --fcc-editable-region--
    
    --fcc-editable-region--
  </table>
</body>
</html>
```
