---
id: 68c33e502cd97f2fd52cb63b
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo workshop imparerai come rendere accessibili gli elementi di una tabella costruendo un programma per una conferenza tecnologica.

Inizia aggiungendo un elemento `h1` con il testo `Tech Conference 2025 Schedule`.

# --hints--

Dovresti avere un elemento `h1`.

```js
const h1 = document.querySelector("h1");
assert.exists(h1);
```

Il tuo elemento `h1` dovrebbe avere il testo `Tech Conference 2025 Schedule`.

```js
const h1 = document.querySelector("h1");
assert.strictEqual(h1?.textContent.trim(), "Tech Conference 2025 Schedule");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
    --fcc-editable-region--
    
    --fcc-editable-region--
</body>
</html>
```
