---
id: 68c34c608cf9e53f5d36b1d5
title: Passaggio 10
challengeType: 0
dashedName: step-10
---

# --description--

Successivamente, aggiungi una terza riga alla tabella. Inizia aggiungendo un altro elemento `tr`. All'interno di quell'elemento `tr`, aggiungi un elemento `th` con un attributo `scope` impostato su `"row"` e il contenuto testuale `11:00 AM`. Poi, sotto quell'elemento `th`, aggiungi un elemento `td` con il contenuto testuale `Break`. 

# --hints--

Dovresti avere un terzo elemento `tr` all'interno dell'elemento `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 3);
```

Dovresti avere un elemento `th` all'interno del terzo elemento `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("th"));
```

Il tuo elemento `th` dovrebbe avere il contenuto testuale `11:00 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.textContent.trim(), "11:00 AM");
```

Il tuo elemento `th` dovrebbe avere un attributo `scope` impostato su `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.getAttribute("scope"), "row");
```

Dovresti avere un elemento `td` all'interno del terzo elemento `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("td"));
```

Il tuo elemento `td` dovrebbe avere il contenuto testuale `Break`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.textContent.trim(), "Break");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

--fcc-editable-region--
      
--fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
