---
id: 68c34ace730ca43e5fa7b8d6
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

Ora è il momento di aggiungere un'altra riga alla tabella.

Inizia aggiungendo un altro elemento `tr`. All'interno di quell'elemento `tr`, aggiungi un elemento `th` con un attributo `scope` impostato su `"row"` e il contenuto testuale `10:00 AM`.

Poi, aggiungi tre elementi `td` con il seguente contenuto testuale:

- `Accessibility Deep Dive`
- `CSS for Beginners`
- `Inclusive Design Principles`

# --hints--

Dovresti avere un secondo elemento `tr` all'interno dell'elemento `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 2);
```

Dovresti avere un elemento `th` all'interno del secondo elemento `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[1].querySelector("th"));
```

Il tuo elemento `th` dovrebbe avere il contenuto testuale `10:00 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.textContent.trim(), "10:00 AM");
```

Il tuo elemento `th` dovrebbe avere un attributo `scope` impostato su `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.getAttribute("scope"), "row");
```

Il tuo secondo elemento `tr` dovrebbe avere tre elementi `td`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[1].querySelectorAll("td"), 3);
```

Il tuo primo elemento `td` dovrebbe avere il contenuto testuale `Accessibility Deep Dive`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[0]?.textContent.trim(), "Accessibility Deep Dive");
```

Il tuo secondo elemento `td` dovrebbe avere il contenuto testuale `CSS for Beginners`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[1]?.textContent.trim(), "CSS for Beginners");
```   

Il tuo terzo elemento `td` dovrebbe avere il contenuto testuale `Inclusive Design Principles`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[2]?.textContent.trim(), "Inclusive Design Principles");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
