---
id: 68c353c4b114aa428c01f260
title: Passaggio 14
challengeType: 0
dashedName: step-14
---

# --description--

L'ultimo passaggio è aggiungere un'altra riga alla tabella.

Inizia aggiungendo un altro elemento `tr`. All'interno di quell'elemento `tr`, aggiungi un elemento `th` con un attributo `scope` impostato su `"row"` e il contenuto testuale `2:00 PM`. 

Poi, aggiungi tre elementi `td` con il seguente contenuto testuale:

- `Voice UI Workshop`
- `Git & GitHub Essentials`
- `Color & Contrast in UI`

Con queste ultime modifiche, la tua tabella è ora completa!

# --hints--

Dovresti avere un sesto elemento `tr` all'interno dell'elemento `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 6);
```

Dovresti avere un elemento `th` all'interno del sesto elemento `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[5].querySelector("th"));
```

Il tuo elemento `th` dovrebbe avere il contenuto testuale `2:00 PM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelector("th")?.textContent.trim(), "2:00 PM");
```

Il tuo elemento `th` dovrebbe avere un attributo `scope` impostato su `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelector("th")?.getAttribute("scope"), "row");
```

Il tuo sesto elemento `tr` dovrebbe avere tre elementi `td`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[5].querySelectorAll("td"), 3);
```

Il tuo primo elemento `td` dovrebbe avere il contenuto testuale `Voice UI Workshop`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelectorAll("td")[0]?.textContent.trim(), "Voice UI Workshop");
```

Il tuo secondo elemento `td` dovrebbe avere il contenuto testuale `Git & GitHub Essentials`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelectorAll("td")[1]?.textContent.trim(), "Git & GitHub Essentials");
```

Il tuo terzo elemento `td` dovrebbe avere il contenuto testuale `Color & Contrast in UI`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelectorAll("td")[2]?.textContent.trim(), "Color & Contrast in UI");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      <tr>
        <th scope="row">12:30 PM</th>
        <td colspan="3">Lunch Break</td>
      </tr>

    --fcc-editable-region--
      
    --fcc-editable-region--
    </tbody>
  </table>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      <tr>
        <th scope="row">12:30 PM</th>
        <td colspan="3">Lunch Break</td>
      </tr>

      <tr>
        <th scope="row">2:00 PM</th>
        <td>Voice UI Workshop</td>
        <td>Git & GitHub Essentials</td>
        <td>Color & Contrast in UI</td>
      </tr>
    </tbody>
  </table>
</body>
</html>
```
