---
id: 68c341332d98fe38fae733af
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Ora è il momento di iniziare a costruire la tabella.

Inizia aggiungendo un elemento `table` sotto il tuo elemento `h1`. All'interno dell'elemento `table`, aggiungi un elemento `caption` con il testo `Schedule by Track and Time`.

# --hints--

Dovresti avere un elemento `table`.

```js
assert.exists(document.querySelector("table"));
```

Il tuo elemento `table` dovrebbe avere un elemento `caption`.

```js
assert.exists(document.querySelector("table caption"));
```

Il tuo elemento `caption` dovrebbe avere il testo `Schedule by Track and Time`.

```js
assert.equal(document.querySelector("table caption")?.textContent.trim(), "Schedule by Track and Time");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

--fcc-editable-region--
    
--fcc-editable-region--
</body>
</html>
```
