---
id: 587d7b8a367417b2b2512b4d
title: استخدام تعيين بالتفكيك لتمرير كائن كمعلمات دالة
challengeType: 1
forumTopicId: 301217
dashedName: use-destructuring-assignment-to-pass-an-object-as-a-functions-parameters
---

# --description--

في بعض الحالات، يمكنك تفكيك الكائن داخل معلمة الدالة نفسها.

انظر الشفرة أدناه:

```js
const profileUpdate = (profileData) => {
  const { name, age, nationality, location } = profileData;

}
```

هذا يُفكك الكائن المرسل إلى الدالة بفعالية. يمكن أيضًا تنفيذ ذلك في المكان نفسه:

```js
const profileUpdate = ({ name, age, nationality, location }) => {

}
```

عند تمرير `profileData` إلى الدالة أعلاه، تُفكك القيم من معلمة الدالة لاستخدامها داخل الدالة.

# --instructions--

استخدم تعيين بالتفكيك داخل المعلمة للدالة `half` لتمرير `max` و `min` فقط داخل الدالة.

# --hints--

يجب أن يكون `stats` عبارة عن `object`.

```js
assert(typeof stats === 'object');
```

يجب أن تكون نتيجة `half(stats)` هي `28.015`

```js
assert(half(stats) === 28.015);
```

يجب استخدام التفكيك.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/half=\({\w+,\w+}\)/));
```

يجب استخدام المعلمة المفككة.

```js
assert(!__helpers.removeJSComments(code).match(/stats\.max|stats\.min/));
```

# --seed--

## --seed-contents--

```js
const stats = {
  max: 56.78,
  standard_deviation: 4.34,
  median: 34.54,
  mode: 23.87,
  min: -0.75,
  average: 35.85
};

// Only change code below this line
const half = (stats) => (stats.max + stats.min) / 2.0; 
// Only change code above this line
```

# --solutions--

```js
const stats = {
  max: 56.78,
  standard_deviation: 4.34,
  median: 34.54,
  mode: 23.87,
  min: -0.75,
  average: 35.85
};

const half = ( {max, min} ) => (max + min) / 2.0;
```
