---
id: 587d7b8b367417b2b2512b50
title: كتابة دوال إعلانية مختصرة باستخدام ES6
challengeType: 1
forumTopicId: 301224
dashedName: write-concise-declarative-functions-with-es6
---

# --description--

عند تعريف الدوال داخل الكائنات في ES5، يجب استخدام الكلمة المفتاحية `function` كما يلي:

```js
const person = {
  name: "Taylor",
  sayHello: function() {
    return `Hello! My name is ${this.name}.`;
  }
};
```

مع ES6، يمكنك حذف الكلمة المفتاحية `function` والنقطتين تمامًا عند تعريف الدوال داخل الكائنات. إليك مثالًا على هذه الصيغة:

```js
const person = {
  name: "Taylor",
  sayHello() {
    return `Hello! My name is ${this.name}.`;
  }
};
```

# --instructions--

أعد صياغة الدالة `setGear` داخل الكائن `bicycle` لاستخدام صيغة الاختصار الموضحة أعلاه.

# --hints--

لا يجب استخدام تعبير الدالة التقليدي.

```js
assert(!__helpers.removeJSComments(code).match(/function/));
```

يجب أن تكون `setGear` دالة إعلانية.

```js
assert(
  typeof bicycle.setGear === 'function' && __helpers.removeJSComments(code).match(/setGear\s*\(.+\)\s*\{/)
);
```

يجب أن تغيّر `bicycle.setGear(48)` قيمة `gear` إلى 48.

```js
bicycle.setGear(48);
assert(bicycle.gear === 48);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const bicycle = {
  gear: 2,
  setGear: function(newGear) {
    this.gear = newGear;
  }
};
// Only change code above this line
bicycle.setGear(3);
console.log(bicycle.gear);
```

# --solutions--

```js
const bicycle = {
  gear: 2,
  // setGear: function(newGear) {
  setGear(newGear) {
    this.gear = newGear;
  }
};
bicycle.setGear(3);
```
