---
id: 587d7b8b367417b2b2512b53
title: استخدام صيغة الصنف لتعريف دالة البناء
challengeType: 1
forumTopicId: 301212
dashedName: use-class-syntax-to-define-a-constructor-function
---

# --description--

يوفر ES6 صيغة جديدة لإنشاء الكائنات باستخدام الكلمة المفتاحية <dfn>class</dfn>.

في ES5، يمكن إنشاء كائن بتعريف دالة `constructor` واستخدام الكلمة المفتاحية `new` لإنشاء نسخة من الكائن.

في ES6، إعلان `class` يحتوي على طريقة `constructor` تُستدعى باستخدام الكلمة المفتاحية `new`. إذا لم تُعرف طريقة `constructor` صراحةً، تُعرف ضمنيًا بدون معلمات.

```js
// Explicit constructor
class SpaceShuttle {
  constructor(targetPlanet) {
    this.targetPlanet = targetPlanet;
  }
  takeOff() {
    console.log("To " + this.targetPlanet + "!");
  }
}

// Implicit constructor 
class Rocket {
  launch() {
    console.log("To the moon!");
  }
}

const zeus = new SpaceShuttle('Jupiter');
// prints To Jupiter! in console
zeus.takeOff();

const atlas = new Rocket();
// prints To the moon! in console
atlas.launch();
```

يجب ملاحظة أن الكلمة المفتاحية `class` تعلن عن دالة جديدة تُضاف إليها دالة البناء. تُستدعى هذه الدالة عند استخدام `new` لإنشاء كائن جديد.

**ملاحظة:** يجب استخدام تنسيق سنام الجمل الكبير (UpperCamelCase) كاصطلاح لأسماء أصناف ES6، كما في `SpaceShuttle` المستخدم أعلاه.

طريقة `constructor` هي طريقة خاصة لإنشاء وتهيئة كائن يُنشأ باستخدام الصنف. ستتعلم المزيد عنها في قسم البرمجة كائنية التوجه في شهادة خوارزميات JavaScript وهياكل البيانات.

# --instructions--

استخدم الكلمة المفتاحية `class` واكتب دالة `constructor` لإنشاء صنف `Vegetable`.

يسمح صنف `Vegetable` بإنشاء كائن خضار مع خاصية `name` تُمرر إلى دالة `constructor`.

# --hints--

يجب أن يكون `Vegetable` عبارة عن `class` مع وجود `constructor` معرف.

```js
assert(
  typeof Vegetable === 'function' && typeof Vegetable.constructor === 'function'
);
```

يجب استخدام الكلمة المفتاحية `class`.

```js
assert(__helpers.removeJSComments(code).match(/class/g));
```

يجب أن يكون من الممكن إنشاء نسخة من `Vegetable`.

```js
assert(() => {
  const a = new Vegetable('apple');
  return typeof a === 'object';
});
```

يجب أن تُرجع `carrot.name` القيمة `carrot`.

```js
assert(carrot.name == 'carrot');
```

# --seed--

## --seed-contents--

```js
// Only change code below this line

// Only change code above this line

const carrot = new Vegetable('carrot');
console.log(carrot.name); // Should display 'carrot'
```

# --solutions--

```js
class Vegetable {
  constructor(name) {
    this.name = name;
  }
}
const carrot = new Vegetable('carrot');
```
