---
id: 587d7b88367417b2b2512b47
title: استخدام معلمة الباقي مع معلمات الدالة
challengeType: 1
forumTopicId: 301221
dashedName: use-the-rest-parameter-with-function-parameters
---

# --description--

لمساعدتنا في إنشاء دوال أكثر مرونة، قدم ES6 <dfn>معلمة الباقي</dfn> لمعلمات الدالة. باستخدام معلمة الباقي، يمكنك إنشاء دوال تأخذ عددًا متغيرًا من المعلمات. تُخزن هذه المعلمات في مصفوفة يمكن الوصول إليها لاحقًا من داخل الدالة.

اطلع على هذه الشفرة:

```js
function howMany(...args) {
  return "You have passed " + args.length + " arguments.";
}
console.log(howMany(0, 1, 2));
console.log(howMany("string", null, [1, 2, 3], { }));
```

ستعرض وحدة التحكم السلاسل النصية `You have passed 3 arguments.` و `You have passed 4 arguments.`.

تلغي معلمة الباقي الحاجة لاستخدام كائن `arguments` وتسمح لنا باستخدام طرق المصفوفة على مصفوفة المعلمات الممررة إلى الدالة `howMany`.

# --instructions--

عدّل الدالة `sum` باستخدام معلمة الباقي بحيث تستطيع الدالة `sum` قبول أي عدد من المعلمات وإرجاع مجموعها.

# --hints--

يجب أن تكون نتيجة `sum(0,1,2)` هي 3

```js
assert(sum(0, 1, 2) === 3);
```

يجب أن تكون نتيجة `sum(1,2,3,4)` هي 10

```js
assert(sum(1, 2, 3, 4) === 10);
```

يجب أن تكون نتيجة `sum(5)` هي 5

```js
assert(sum(5) === 5);
```

يجب أن تكون نتيجة `sum()` هي 0

```js
assert(sum() === 0);
```

يجب أن تكون `sum` دالة سهمية تستخدم صيغة معلمة الباقي (`...`) على المعلمة `args`.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/sum=\(\.\.\.args\)=>/));
```

# --seed--

## --seed-contents--

```js
const sum = (x, y, z) => {
  const args = [x, y, z];
  let total = 0;
  for (let i = 0; i < args.length; i++) {
    total += args[i];
  }
  return total;
}
```

# --solutions--

```js
const sum = (...args) => {
  let total = 0;
  for (let i = 0; i < args.length; i++) {
    total += args[i];
  }
  return total;
}
```
