---
id: 587d7b8c367417b2b2512b56
title: استخدام export لمشاركة كتلة شفرة
challengeType: 1
forumTopicId: 301219
dashedName: use-export-to-share-a-code-block
---

# --description--

تخيل ملفًا اسمه `math_functions.js` يحتوي على عدة دوال مرتبطة بالعمليات الرياضية. واحدة منها مخزنة في متغير، `add`، يأخذ رقمين ويُرجع مجموعهما. تريد استخدام هذه الدالة في عدة ملفات JavaScript مختلفة. لمشاركتها مع هذه الملفات الأخرى، يجب أولًا أن تقوم بـ `export` لها.

```js
export const add = (x, y) => {
  return x + y;
}
```

الطريقة السابقة شائعة لتصدير دالة واحدة، لكن يمكنك تحقيق نفس الشيء بهذه الطريقة:

```js
const add = (x, y) => {
  return x + y;
}

export { add };
```

عندما تقوم بتصدير متغير أو دالة، يمكنك استيرادها في ملف آخر واستخدامها دون الحاجة لإعادة كتابة الشفرة. يمكنك تصدير عدة أشياء بتكرار المثال الأول لكل شيء تريد تصديره، أو بوضعهم جميعًا في جملة التصدير في المثال الثاني، مثل هذا:

```js
export { add, subtract };
```

# --instructions--

هناك دالتان مرتبطتان بالسلاسل النصية في المحرر. صدّر كلاهما باستخدام الطريقة التي تختارها.

# --hints--

يجب أن تصدر `uppercaseString` بشكل صحيح.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(export\s+const\s+uppercaseString|export\s*{\s*(uppercaseString[^}]*|[^,]*,\s*uppercaseString\s*)})/g
  )
);
```

يجب أن تصدر `lowercaseString` بشكل صحيح.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(export\s+const\s+lowercaseString|export\s*{\s*(lowercaseString[^}]*|[^,]*,\s*lowercaseString\s*)})/g
  )
);
```

# --seed--

## --seed-contents--

```js
const uppercaseString = (string) => {
  return string.toUpperCase();
}

const lowercaseString = (string) => {
  return string.toLowerCase()
}
```

# --solutions--

```js
export const uppercaseString = (string) => {
  return string.toUpperCase();
}

export const lowercaseString = (string) => {
  return string.toLowerCase()
}
```
