---
id: 587d7b8a367417b2b2512b4c
title: "التعيين بالتفكيك عبر عناصر rest"
challengeType: 1
forumTopicId: 301218
dashedName: >-
  use-destructuring-assignment-with-the-rest-parameter-to-reassign-array-elements
---

# --description--

في بعض الحالات التي تتضمن التفكيك لمصفوفة، قد نرغب في جمع بقية العناصر في مصفوفة منفصلة.

النتيجة مشابهة لدالة `Array.prototype.slice()`، كما هو موضح أدناه:

```js
const [a, b, ...arr] = [1, 2, 3, 4, 5, 7];
console.log(a, b);
console.log(arr);
```

ستعرض وحدة التحكم القيم `1, 2` و `[3, 4, 5, 7]`.

المتغيران `a` و `b` يأخذان القيمة الأولى والثانية من المصفوفة. بعد ذلك، بسبب وجود صيغة rest، يحصل `arr` على بقية القيم في شكل مصفوفة. عنصر rest يعمل بشكل صحيح فقط كآخر متغير في القائمة. بمعنى آخر، لا يمكنك استخدام صيغة rest لالتقاط مصفوفة فرعية تستثني العنصر الأخير من المصفوفة الأصلية.

# --instructions--

استخدم تعيينًا بالتفكيك مع صيغة rest لمحاكاة سلوك `Array.prototype.slice()`. يجب أن تُرجع الدالة `removeFirstTwo()` مصفوفة فرعية من المصفوفة الأصلية `list` مع حذف أول عنصرين.

# --hints--

`removeFirstTwo([1, 2, 3, 4, 5])` يجب أن تكون `[3, 4, 5]`

```js
assert.deepEqual(removeFirstTwo([1, 2, 3, 4, 5]), [3, 4, 5]);
```

يجب ألا تُعدّل `removeFirstTwo()` المصفوفة `list`

```js
const _testArr = [1, 2, 3, 4, 5];
removeFirstTwo(_testArr);
assert.deepEqual(_testArr, [1, 2, 3, 4, 5])
```

يجب عدم استخدام `Array.slice()`.

```js
assert(!__helpers.removeJSComments(code).match(/\.\s*slice\s*\(/));
```

يجب أن تستخدم صيغة rest.

```js
assert.match(code, /\.\.\./);
```

# --seed--

## --seed-contents--

```js
function removeFirstTwo(list) {
  return list;
}

const source = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const sourceWithoutFirstTwo = removeFirstTwo(source);
```

# --solutions--

```js
function removeFirstTwo(list) {
  // comment with 'slice' to check comments are removed in tests
  const [, , ...shorterList] = list;
  return shorterList;
}

const source = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const sourceWithoutFirstTwo = removeFirstTwo(source);
```
