---
id: 587d7b8c367417b2b2512b57
title: استخدام * لاستيراد كل شيء من ملف
challengeType: 1
forumTopicId: 301210
dashedName: use--to-import-everything-from-a-file
---

# --description--

افترض أن لديك ملفًا وترغب في استيراد كل محتوياته إلى الملف الحالي. يمكن فعل ذلك باستخدام صيغة `import * as`. إليك مثالًا حيث تُستورد محتويات ملف اسمه `math_functions.js` إلى ملف في نفس المجلد:

```js
import * as myMathModule from "./math_functions.js";
```

عبارة `import` أعلاه ستُنشئ كائنًا اسمه `myMathModule`. هذا مجرد اسم متغير، يمكنك تسميته بأي اسم تريده. سيحتوي الكائن على كل الصادرات من `math_functions.js`، لذا يمكنك الوصول إلى الدوال كما تفعل مع أي خاصية في كائن آخر. إليك كيف تستخدم دالتي `add` و `subtract` اللتين تم استيرادهما:

```js
myMathModule.add(2,3);
myMathModule.subtract(5,3);
```

# --instructions--

الكود في هذا الملف يحتاج إلى محتويات الملف: `string_functions.js`، الموجود في نفس المجلد مع الملف الحالي. استخدم صيغة `import * as` لاستيراد كل شيء من الملف إلى كائن اسمه `stringFunctions`.

# --hints--

يجب أن يستخدم كودك صيغة `import * as` بشكل صحيح.

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*\*\s*as\s+stringFunctions\s+from\s*('|")\.\/string_functions\.js\1/g
  )
);
```

# --seed--

## --seed-contents--

```js

// Only change code above this line

stringFunctions.uppercaseString("hello");
stringFunctions.lowercaseString("WORLD!");
```

# --solutions--

```js
import * as stringFunctions from "./string_functions.js";

// add code above this line
stringFunctions.uppercaseString("hello");
stringFunctions.lowercaseString("WORLD!");
```
