---
id: 587d7b89367417b2b2512b4b
title: استخدام تعيين بالتفكيك لتعيين المتغيرات من المصفوفات
challengeType: 1
forumTopicId: 301213
dashedName: use-destructuring-assignment-to-assign-variables-from-arrays
---

# --description--

يُسهّل ES6 تفكيك المصفوفات بنفس سهولة تفكيك الكائنات.

الفرق الرئيسي بين معامل النشر والتفكيك في المصفوفات هو أن معامل النشر يفكّ كل محتويات المصفوفة إلى قائمة مفصولة بفواصل. لذلك، لا يمكنك اختيار العناصر التي تريد تعيينها إلى متغيرات.

يسمح لنا تفكيك المصفوفة بفعل ذلك بالضبط:

```js
const [a, b] = [1, 2, 3, 4, 5, 6];
console.log(a, b);
```

ستعرض وحدة التحكم قيم `a` و `b` كـ `1, 2`.

يُعيَّن المتغير `a` إلى القيمة الأولى في المصفوفة، و`b` يُعيَّن إلى القيمة الثانية في المصفوفة. يمكننا أيضًا الوصول إلى القيمة في أي فهرس داخل المصفوفة باستخدام التفكيك عبر استخدام الفواصل للوصول إلى الفهرس المطلوب:

```js
const [a, b,,, c] = [1, 2, 3, 4, 5, 6];
console.log(a, b, c);
```

ستعرض وحدة التحكم قيم `a` و `b` و `c` كـ `1, 2, 5`.

# --instructions--

استخدم تعيين بالتفكيك لتبديل قيم `a` و `b` بحيث يستقبل `a` القيمة المخزنة في `b`، ويستقبل `b` القيمة المخزنة في `a`.

# --hints--

يجب أن تكون قيمة `a` هي `6` بعد التبديل.

```js
assert(a === 6);
```

يجب أن تكون قيمة `b` هي `8` بعد التبديل.

```js
assert(b === 8);
```

يجب أن تستخدم تفكيك المصفوفة لتبديل `a` و `b`.

```js
assert(/\[\s*(\w)\s*,\s*(\w)\s*\]\s*=\s*\[\s*\2\s*,\s*\1\s*\]/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let a = 8, b = 6;
// Only change code below this line
```

# --solutions--

```js
let a = 8, b = 6;
[a, b] = [b, a];
```
