---
id: 587d7b87367417b2b2512b40
title: مقارنة نطاقات الكلمات المفتاحية var و let
challengeType: 1
forumTopicId: 301195
dashedName: compare-scopes-of-the-var-and-let-keywords
---

# --description--

إذا لم تكن معتادًا على `let`، اطلع على <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/explore-differences-between-the-var-and-let-keywords" target="_blank" rel="noopener noreferrer nofollow">هذا التحدٍّ حول الفرق بين <code>let</code> و <code>var</code></a>.

عندما تعلن عن متغير باستخدام الكلمة المفتاحية `var`، يُعلن عنه عالميًا، أو محليًا إذا أُعلن داخل دالة.

تتصرف الكلمة المفتاحية `let` بطريقة مشابهة، لكنها تضيف بعض الميزات. عندما تعلن عن متغير باستخدام `let` داخل كتلة أو تعبير أو جملة، يقتصر نطاقه على تلك الكتلة أو الجملة أو التعبير فقط.

على سبيل المثال:

```js
var numArray = [];
for (var i = 0; i < 3; i++) {
  numArray.push(i);
}
console.log(numArray);
console.log(i);
```

هنا ستعرض وحدة التحكم القيم `[0, 1, 2]` و `3`.

مع الكلمة المفتاحية `var`، يُعلن عن `i` عالميًا. لذا عند تنفيذ `i++`، يتم تحديث المتغير العام. هذا الكود مشابه لما يلي:

```js
var numArray = [];
var i;
for (i = 0; i < 3; i++) {
  numArray.push(i);
}
console.log(numArray);
console.log(i);
```

هنا ستعرض وحدة التحكم القيم `[0, 1, 2]` و `3`.

هذا السلوك سيسبب مشاكل إذا أنشأت دالة وخزنتها للاستخدام لاحقًا داخل حلقة `for` تستخدم المتغير `i`. لأن الدالة المخزنة ستشير دائمًا إلى قيمة المتغير العام `i` المحدث.

```js
var printNumTwo;
for (var i = 0; i < 3; i++) {
  if (i === 2) {
    printNumTwo = function() {
      return i;
    };
  }
}
console.log(printNumTwo());
```

هنا ستعرض وحدة التحكم القيمة `3`.

كما ترى، `printNumTwo()` تطبع 3 وليس 2. هذا لأن القيمة المعينة لـ `i` تم تحديثها و`printNumTwo()` تعيد المتغير العام `i` وليس القيمة التي كان عليها `i` عند إنشاء الدالة داخل حلقة for. الكلمة المفتاحية `let` لا تتبع هذا السلوك:

```js
let printNumTwo;
for (let i = 0; i < 3; i++) {
  if (i === 2) {
    printNumTwo = function() {
      return i;
    };
  }
}
console.log(printNumTwo());
console.log(i);
```

هنا ستعرض وحدة التحكم القيمة `2`، وخطأ يفيد بأن `i is not defined`.

لم يُعرف `i` لأنه لم يُعلن في النطاق العام. بل أُعلن فقط داخل جملة حلقة `for`. أعادت `printNumTwo()` القيمة الصحيحة لأن الكلمة المفتاحية `let` أنشأت ثلاثة متغيرات `i` مختلفة بقيم فريدة (0، 1، و2) داخل جملة الحلقة.

# --instructions--

صحح الشفرة بحيث يكون المتغير `i` المعلن في جملة `if` متغيرًا منفصلًا عن `i` المعلن في السطر الأول من الدالة. تأكد من عدم استخدام الكلمة المفتاحية `var` في أي مكان في شفرتك.

هذا التمرين مصمم لتوضيح الفرق بين كيفية تعيين الكلمات المفتاحية `var` و `let` للنطاق للمتغير المعلن. عند برمجة دالة مشابهة لتلك المستخدمة في هذا التمرين، من الأفضل غالبًا استخدام أسماء متغيرات مختلفة لتجنب الالتباس.

# --hints--

يجب ألا توجد كلمة `var` في الشفرة.

```js
assert(!__helpers.removeJSComments(code).match(/var/g));
```

يجب أن يساوي المتغير `i` المعلن في جملة `if` السلسلة النصية `block scope`.

```js
assert(__helpers.removeJSComments(code).match(/(i\s*=\s*).*\s*.*\s*.*\1('|")block\s*scope\2/g));
```

يجب أن تُعيد الدالة `checkScope()` السلسلة النصية `function scope`.

```js
assert(checkScope() === 'function scope');
```

# --seed--

## --seed-contents--

```js
function checkScope() {
  var i = 'function scope';
  if (true) {
    i = 'block scope';
    console.log('Block scope i is: ', i);
  }
  console.log('Function scope i is: ', i);
  return i;
}
```

# --solutions--

```js
function checkScope() {
  let i = 'function scope';
  if (true) {
    let i = 'block scope';
    console.log('Block scope i is: ', i);
  }
 
  console.log('Function scope i is: ', i);
  return i;
}
```
