---
id: 587d7b87367417b2b2512b43
title: استخدام الدوال السهمية لكتابة دوال مجهولة مختصرة
challengeType: 1
forumTopicId: 301211
dashedName: use-arrow-functions-to-write-concise-anonymous-functions
---

# --description--

في JavaScript، غالبًا لا نحتاج إلى تسمية دوالنا، خاصة عند تمرير دالة كمعلمة إلى دالة أخرى. بدلاً من ذلك، ننشئ دوالًا مضمَّنة. لا نحتاج إلى تسمية هذه الدوال لأننا لا نُعيد استخدامها في مكان آخر.

لتحقيق ذلك، نستخدم غالبًا الصيغة التالية:

```js
const myFunc = function() {
  const myVar = "value";
  return myVar;
}
```

يوفر لنا ES6 صيغة نحوية مبسطة لعدم الاضطرار إلى كتابة دوال مجهولة بهذه الطريقة. بدلاً من ذلك، يمكنك استخدام **صيغة الدالة السهمية**:

```js
const myFunc = () => {
  const myVar = "value";
  return myVar;
}
```

عندما لا يكون هناك جسم للدالة، ويوجد فقط قيمة إرجاع، تسمح صيغة الدالة السهمية بحذف الكلمة المفتاحية `return` وكذلك الأقواس المحيطة بالشفرة. هذا يساعد على تبسيط الدوال الصغيرة إلى عبارات في سطر واحد:

```js
const myFunc = () => "value";
```

ستُرجع هذه الشفرة سلسلة نصية `value` افتراضيًا.

# --instructions--

أعد كتابة الدالة المُعيَّنة للمتغير `magic` والتي تُرجع `new Date()` باستخدام صيغة الدالة السهمية. تأكد أيضًا من عدم تعريف أي شيء باستخدام الكلمة المفتاحية `var`.

# --hints--

يجب استبدال الكلمة المفتاحية `var`.

```js
assert.notMatch(code, /var/g)
```

يجب أن يكون `magic` متغيرًا ثابتًا (باستخدام `const`).

```js
assert.match(code, /const\s+magic/g)
```

يجب أن تكون `magic` دالة `function`.

```js
assert(typeof magic === 'function');
```

يجب أن تُرجع `magic()` التاريخ الصحيح.

```js
assert(magic().setHours(0, 0, 0, 0) === new Date().setHours(0, 0, 0, 0));
```

يجب عدم استخدام الكلمة المفتاحية `function`.

```js
assert.notMatch(code, /function/g)
```

# --seed--

## --seed-contents--

```js
var magic = function() {
  return new Date();
};
```

# --solutions--

```js
const magic = () => {
  return new Date();
};
```
