---
id: 587d7b8c367417b2b2512b54
title: استخدام getters و setters للتحكم في الوصول إلى كائن
challengeType: 1
forumTopicId: 301220
dashedName: use-getters-and-setters-to-control-access-to-an-object
---

# --description--

يمكنك الحصول على قيم من كائن وتعيين قيمة خاصية داخل الكائن.

تُسمى هذه تقليديًا <dfn>getters</dfn> و <dfn>setters</dfn>.

تهدف دوال getter إلى إرجاع (الحصول على) قيمة المتغير الخاص في الكائن للمستخدم دون أن يصل المستخدم مباشرة إلى المتغير الخاص.

تهدف دوال setter إلى تعديل (تعيين) قيمة المتغير الخاص في الكائن بناءً على القيمة التي تُمرر إلى دالة setter. قد يشمل هذا التغيير حسابات، أو حتى استبدال القيمة السابقة بالكامل.

```js
class Book {
  constructor(author) {
    this._author = author;
  }
  // getter
  get writer() {
    return this._author;
  }
  // setter
  set writer(updatedAuthor) {
    this._author = updatedAuthor;
  }
}
const novel = new Book('anonymous');
console.log(novel.writer);
novel.writer = 'newAuthor';
console.log(novel.writer);
```

ستعرض وحدة التحكم السلاسل النصية `anonymous` و `newAuthor`.

لاحظ الصيغة النحوية المستخدمة لاستدعاء getter و setter. لا تبدو حتى كدوال. getters و setters مهمة لأنها تخفي تفاصيل التنفيذ الداخلية.

**ملاحظة:** من الاصطلاح أن يسبق اسم المتغير الخاص بشرطة سفلية (`_`). مع ذلك، هذه الممارسة لا تجعل المتغير خاصًا فعليًا.

# --instructions--

استخدم الكلمة المفتاحية `class` لإنشاء صنف `Thermostat`. تستقبل دالة البناء `constructor` درجة حرارة بالفهرنهايت.

في الصنف، أنشئ `getter` للحصول على درجة الحرارة بالسلسيوس و `setter` يقبل درجة حرارة بالسلسيوس.

تذكر أن `C = 5/9 * (F - 32)` و `F = C * 9.0 / 5 + 32`، حيث `F` هي قيمة درجة الحرارة بالفهرنهايت، و `C` هي قيمة نفس درجة الحرارة بالسلسيوس.

**ملاحظة:** عند تنفيذ هذا، ستتتبع درجة الحرارة داخل الصنف بمقياس واحد، إما الفهرنهايت أو السلسيوس.

هذه هي قوة getter و setter. أنت تخلق واجهة برمجة تطبيقات (API) لمستخدم آخر، يمكنه الحصول على النتيجة الصحيحة بغض النظر عن المقياس الذي تتتبعه.

بمعنى آخر، أنت تُجرد تفاصيل التنفيذ عن المستخدم.

# --hints--

يجب أن يكون `Thermostat` عبارة عن `class` مع تعريف `constructor`.

```js
assert.isFunction(Thermostat);
assert.isFunction(Thermostat?.constructor);
```

يجب استخدام الكلمة المفتاحية `class`.

```js
assert.match(code, /class/);
```

يجب أن يكون بالإمكان إنشاء نسخة من `Thermostat`.

```js
const _t = new Thermostat(122);
assert.isObject(_t);
```

عند إنشاء كائن من `Thermostat` بقيمة فهرنهايت، يجب أن يضبط `temperature` بشكل صحيح.

```js
const _t = new Thermostat(122);
assert.strictEqual(_t?.temperature, 50);
```

يجب تعريف `getter`.

```js
const _desc = Object.getOwnPropertyDescriptor(Thermostat.prototype, 'temperature');
assert.isFunction(_desc?.get);
```

يجب تعريف `setter`.

```js
const _desc = Object.getOwnPropertyDescriptor(Thermostat.prototype, 'temperature');
assert.isFunction(_desc?.set);
```

استدعاء الـ `setter` بقيمة مئوية يجب أن يضبط الـ `temperature`.

```js
const _t = new Thermostat(32);
_t.temperature = 26;
const _u = new Thermostat(32);
_u.temperature = 50;
assert.approximately(_t.temperature, 26, 0.1);
assert.approximately(_u.temperature, 50, 0.1);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line

// Only change code above this line

const thermos = new Thermostat(76); // Setting in Fahrenheit scale
let temp = thermos.temperature; // 24.44 in Celsius
thermos.temperature = 26;
temp = thermos.temperature; // 26 in Celsius
```

# --solutions--

```js
class Thermostat {
  constructor(fahrenheit) {
    this._tempInCelsius = 5/9 * (fahrenheit - 32);
  }
  get temperature(){
    return this._tempInCelsius;
  }
  set temperature(newTemp){
    this._tempInCelsius = newTemp;
  }
}

const thermos = new Thermostat(76); // Setting in Fahrenheit scale
let temp = thermos.temperature; // 24.44 in Celsius
thermos.temperature = 26;
temp = thermos.temperature; // 26 in Celsius
```
