---
id: 587d7b89367417b2b2512b4a
title: استخدام تعيين بالتفكيك لتعيين المتغيرات من كائنات مسكنة
challengeType: 1
forumTopicId: 301214
dashedName: use-destructuring-assignment-to-assign-variables-from-nested-objects
---

# --description--

يمكنك استخدام نفس مبادئ الدرسين السابقين لتفكيك القيم من كائنات مسكنة.

باستخدام كائن مشابه للأمثلة السابقة:

```js
const user = {
  johnDoe: { 
    age: 34,
    email: 'johnDoe@freeCodeCamp.com'
  }
};
```

إليك كيفية استخراج قيم خصائص الكائن وتعيينها إلى متغيرات بنفس الاسم:

```js
const { johnDoe: { age, email }} = user;
```

وهنا كيفية تعيين قيم خصائص الكائن إلى متغيرات بأسماء مختلفة:

```js
const { johnDoe: { age: userAge, email: userEmail }} = user;
```

# --instructions--

استبدل التعيينين بتعيين بالتفكيك مكافئ. يجب أن يظل يعين المتغيرين `lowToday` و `highToday` القيم من `today.low` و `today.high` من كائن `LOCAL_FORECAST`.

# --hints--

يجب إزالة صيغة التعيين في ES5.

```js
assert(
  !__helpers.removeJSComments(code).match(/lowToday = LOCAL_FORECAST\.today\.low/g) &&
    !__helpers.removeJSComments(code).match(/highToday = LOCAL_FORECAST\.today.high/g)
);
```

يجب استخدام التفكيك لإنشاء المتغير `lowToday`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*today\s*:\s*{\s*(low\s*:\s*lowToday[^}]*|[^,]*,\s*low\s*:\s*lowToday\s*)},?\s*}\s*=\s*LOCAL_FORECAST(;|\s+|\/\/)/g
  )
);
```

يجب استخدام التفكيك لإنشاء المتغير `highToday`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*today\s*:\s*{\s*(high\s*:\s*highToday[^}]*|[^,]*,\s*high\s*:\s*highToday,?\s*)},?\s*}\s*=\s*LOCAL_FORECAST(;|\s+|\/\/)/g
  )
);
```

يجب أن تكون قيمة `lowToday` مساوية لـ `64` وقيمة `highToday` مساوية لـ `77`.

```js
assert(lowToday === 64 && highToday === 77);
```

# --seed--

## --seed-contents--

```js
const LOCAL_FORECAST = {
  yesterday: { low: 61, high: 75 },
  today: { low: 64, high: 77 },
  tomorrow: { low: 68, high: 80 }
};

// Only change code below this line
  
const lowToday = LOCAL_FORECAST.today.low;
const highToday = LOCAL_FORECAST.today.high;

// Only change code above this line
```

# --solutions--

```js
const LOCAL_FORECAST = {
  yesterday: { low: 61, high: 75 },
  today: { low: 64, high: 77 },
  tomorrow: { low: 68, high: 80 }
};
 
const { today: { low: lowToday, high: highToday }} = LOCAL_FORECAST;
```
