---
id: 587d7b88367417b2b2512b44
title: كتابة الدوال السهمية مع المعلمات
challengeType: 1
forumTopicId: 301223
dashedName: write-arrow-functions-with-parameters
---

# --description--

تمامًا مثل الدالة العادية، يمكنك تمرير المعلمات إلى الدالة السهمية.

```js
const doubler = (item) => item * 2;
doubler(4);
```

`doubler(4)` سترجع القيمة `8`.

إذا كانت الدالة السهمية تحتوي على معلمة واحدة، يمكن حذف الأقواس العادية التي تحيط بالمعلمة.

```js
const doubler = item => item * 2;
```

يمكن تمرير أكثر من معلمة إلى الدالة السهمية.

```js
const multiplier = (item, multi) => item * multi;
multiplier(4, 2);
```

`multiplier(4, 2)` سترجع القيمة `8`.

# --instructions--

أعد كتابة الدالة `myConcat` التي تلحق محتويات `arr2` إلى `arr1` بحيث تستخدم الدالة صيغة الدالة السهمية.

# --hints--

يجب استبدال الكلمة المفتاحية `var`.

```js
assert.notMatch(code, /var/g);
```

يجب أن تكون `myConcat` متغيرًا ثابتًا (باستخدام `const`).

```js
assert.match(code, /const\s+myConcat/g);
```

يجب أن تكون `myConcat` دالة سهمية تحتوي على معلمتين

```js
assert(
  /myConcat=\(\w+,\w+\)=>/.test(__helpers.removeJSComments(code).replace(/\s/g, '')) &&
    typeof myConcat === 'function'
);
```

`myConcat()` يجب أن ترجع `[1, 2, 3, 4, 5]`.

```js
assert.deepEqual(myConcat([1, 2], [3, 4, 5]), [1, 2, 3, 4, 5]);
```

يجب عدم استخدام الكلمة المفتاحية `function`.

```js
assert.notMatch(code, /function/g);
```

# --seed--

## --seed-contents--

```js
var myConcat = function(arr1, arr2) {
  return arr1.concat(arr2);
};

console.log(myConcat([1, 2], [3, 4, 5]));
```

# --solutions--

```js
const myConcat = (arr1, arr2) =>  {
  return arr1.concat(arr2);
};

console.log(myConcat([1, 2], [3, 4, 5]));
```
