---
id: 587d7b8c367417b2b2512b58
title: إنشاء نظام بديل للتصدير باستخدام export default
challengeType: 1
forumTopicId: 301199
dashedName: create-an-export-fallback-with-export-default
---

# --description--

في درس `export`، تعلمت عن الصيغة النحوية المعروفة باسم <dfn>التصدير المسمي</dfn>. سمحت لك هذه الصيغة بجعل عدة دوال ومتغيرات متاحة للاستخدام في ملفات أخرى.

هناك صيغة أخرى لـ `export` يجب أن تعرفها، تُسمى <dfn>export default</dfn>. عادةً تستخدم هذه الصيغة إذا كان يتم تصدير قيمة واحدة فقط من ملف. كما تُستخدم لإنشاء قيمة نظام بديل لملف أو وحدة.

فيما يلي أمثلة تستخدم `export default`:

```js
export default function add(x, y) {
  return x + y;
}

export default function(x, y) {
  return x + y;
}
```

الأول هو دالة مسماة، والثاني هو دالة مجهولة.

بما أن `export default` تُستخدم لإعلان قيمة نظام بديل لوحدة أو ملف، يمكنك أن يكون لديك قيمة واحدة فقط كتصدير افتراضي في كل وحدة أو ملف. بالإضافة إلى ذلك، لا يمكنك استخدام `export default` مع `var` أو `let` أو `const`

# --instructions--

يجب أن تكون الدالة التالية هي قيمة النظام البديل للوحدة. أضف الشفرة اللازمة لذلك.

# --hints--

يجب أن يستخدم شفرتك نظام بديل لـ `export`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /export\s+default\s+function(\s+subtract\s*|\s*)\(\s*x,\s*y\s*\)\s*{/g
  )
);
```

# --seed--

## --seed-contents--

```js
function subtract(x, y) {
  return x - y;
}
```

# --solutions--

```js
export default function subtract(x, y) {
  return x - y;
}
```
