---
id: 587d7b8c367417b2b2512b55
title: إعادة استخدام شفرة JavaScript باستخدام import
challengeType: 1
forumTopicId: 301208
dashedName: reuse-javascript-code-using-import
---

# --description--

`import` يسمح لك باختيار أي أجزاء من ملف أو وحدة تحميلها. في الدرس السابق، الأمثلة صدرت `add` من ملف `math_functions.js`. إليك كيف تستوردها لاستخدامها في ملف آخر:

```js
import { add } from './math_functions.js';
```

هنا، `import` سيجد `add` في `math_functions.js`، يستورد تلك الدالة فقط لاستخدامك، ويتجاهل الباقي. الـ `./` يخبر الاستيراد بالبحث عن ملف `math_functions.js` في نفس المجلد الذي فيه الملف الحالي. مسار الملف النسبي (`./`) وامتداد الملف (`.js`) مطلوبان عند استخدام import بهذه الطريقة.

يمكنك استيراد أكثر من عنصر من الملف بإضافتها في جملة `import` هكذا:

```js
import { add, subtract } from './math_functions.js';
```

# --instructions--

أضف جملة `import` المناسبة التي تسمح للملف الحالي باستخدام دالتي `uppercaseString` و `lowercaseString` التي صدرتها في الدرس السابق. هذه الدوال في ملف يسمى `string_functions.js`، وهو في نفس الدليل مع الملف الحالي.

# --hints--

يجب استيراد `uppercaseString` بشكل صحيح.

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*{\s*(uppercaseString[^}]*|[^,]*,\s*uppercaseString\s*)}\s+from\s+('|")\.\/string_functions\.js\2/g
  )
);
```

يجب استيراد `lowercaseString` بشكل صحيح.

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*{\s*(lowercaseString[^}]*|[^,]*,\s*lowercaseString\s*)}\s+from\s+('|")\.\/string_functions\.js\2/g
  )
);
```

# --seed--

## --seed-contents--

```js
  
// Only change code above this line

uppercaseString("hello");
lowercaseString("WORLD!");
```

# --solutions--

```js
import { uppercaseString, lowercaseString } from './string_functions.js';

uppercaseString("hello");
lowercaseString("WORLD!");
```
