---
id: 587d7b8a367417b2b2512b4e
title: إنشاء سلاسل نصية باستخدام قوالب حرفية
challengeType: 1
forumTopicId: 301200
dashedName: create-strings-using-template-literals
---

# --description--

ميزة جديدة في ES6 هي <dfn>القالب الحرفي</dfn>. هذا نوع خاص من السلاسل النصية يجعل إنشاء السلاسل المعقدة أسهل.

تسمح القوالب الحرفية بإنشاء سلاسل نصية متعددة الأسطر واستخدام ميزات الاستيفاء لإنشاء السلاسل.

انظر الكود أدناه:

```js
const person = {
  name: "Zodiac Hasbro",
  age: 56
};

const greeting = `Hello, my name is ${person.name}!
I am ${person.age} years old.`;

console.log(greeting);
```

ستعرض وحدة التحكم السلاسل النصية `Hello, my name is Zodiac Hasbro!` و `I am 56 years old.`.

حدثت عدة أشياء هناك. أولًا، المثال يستخدم علامات دقيقة (`` ` ``) وليس علامات اقتباس (`'` أو `"`) لتغليف السلسلة النصية. ثانيًا، لاحظ أن السلسلة متعددة الأسطر، سواء في الكود أو في الناتج. هذا يوفر عناء إدخال `\n` داخل السلاسل النصية. الصيغة `${variable}` المستخدمة أعلاه هي عنصر نائب. باختصار، لن تحتاج إلى استخدام الجمع مع العامل `+` بعد الآن. لإضافة متغيرات إلى السلاسل النصية، فقط ضع المتغير داخل سلسلة قالب واغلفه بـ `${` و `}`. وبالمثل، يمكنك تضمين تعبيرات أخرى داخل سلسلة القالب، مثل `${a + b}`. هذه الطريقة الجديدة لإنشاء السلاسل تمنحك مرونة أكبر لإنشاء سلاسل قوية.

# --instructions--

استخدم صيغة القالب الحرفي مع علامات دقيقة لإنشاء مصفوفة من سلاسل عناصر القائمة (`li`). يجب أن يكون نص كل عنصر قائمة واحدًا من عناصر المصفوفة في خاصية `failure` على كائن `result`، ويحتوي على خاصية `class` بقيمة `text-warning`. يجب أن تُرجع الدالة `makeList` مصفوفة سلاسل عناصر القائمة.

استخدم طريقة تكرار (أي نوع من الحلقات) للحصول على الإخراج المطلوب (الموضح أدناه).

```js
[
  '<li class="text-warning">no-var</li>',
  '<li class="text-warning">var-on-top</li>',
  '<li class="text-warning">linebreak</li>'
]
```

# --hints--

`failuresList` يجب أن تكون مصفوفة تحتوي على رسائل `result failure`.

```js
assert(
  typeof makeList(result.failure) === 'object' && failuresList.length === 3
);
```

`failuresList` يجب أن تساوي الإخراج المحدد.

```js
assert(
  makeList(result.failure).every(
    (v, i) =>
      v === `<li class="text-warning">${result.failure[i]}</li>` ||
      v === `<li class='text-warning'>${result.failure[i]}</li>`
  )
);
```

يجب استخدام قوالب السلاسل النصية واستيفاء التعبيرات.

```js
assert.match(code, /(`.*\${.*}.*`)/);
```

يجب استخدام طريقة تكرار.

```js
assert(__helpers.removeJSComments(code).match(/for|map|reduce|forEach|while/));
```

# --seed--

## --seed-contents--

```js
const result = {
  success: ["max-length", "no-amd", "prefer-arrow-functions"],
  failure: ["no-var", "var-on-top", "linebreak"],
  skipped: ["no-extra-semi", "no-dup-keys"]
};
function makeList(arr) {
  // Only change code below this line
  const failureItems = [];
  // Only change code above this line

  return failureItems;
}

const failuresList = makeList(result.failure);
```

# --solutions--

```js
const result = {
  success: ["max-length", "no-amd", "prefer-arrow-functions"],
  failure: ["no-var", "var-on-top", "linebreak"],
  skipped: ["no-extra-semi", "no-dup-keys"]
};
function makeList(arr) {
  return arr.map(val => `<li class="text-warning">${val}</li>`);
}

const failuresList = makeList(result.failure);
```
