---
id: 587d7b8a367417b2b2512b4f
title: كتابة إعلانات مختصرة لكائنات الحرف باستخدام اختصار خاصية الكائن
challengeType: 1
forumTopicId: 301225
dashedName: write-concise-object-literal-declarations-using-object-property-shorthand
---

# --description--

يضيف ES6 دعمًا جيدًا لتعريف كائنات الحرف بسهولة.

انظر الشفرة التالية:

```js
const getMousePosition = (x, y) => ({
  x: x,
  y: y
});
```

`getMousePosition` دالة بسيطة تُرجع كائنًا يحتوي على خاصيتين. يوفر ES6 صيغة نحوية لتجنب التكرار عند كتابة `x: x`. يمكنك ببساطة كتابة `x` مرة واحدة، وسيتم تحويلها إلى `x: x` (أو ما يعادلها) ضمنيًا. هذه هي نفس الدالة أعلاه معاد كتابتها باستخدام هذه الصيغة الجديدة:

```js
const getMousePosition = (x, y) => ({ x, y });
```

# --instructions--

استخدم اختصار خاصية الكائن مع كائنات الحرف لإنشاء وإرجاع كائن يحتوي على الخصائص `name` و `age` و `gender`.

# --hints--

`createPerson("Zodiac Hasbro", 56, "male")` يجب أن تُرجع `{name: "Zodiac Hasbro", age: 56, gender: "male"}`.

```js
assert.deepEqual(
  { name: 'Zodiac Hasbro', age: 56, gender: 'male' },
  createPerson('Zodiac Hasbro', 56, 'male')
);
```

يجب ألا تستخدم شفرتك `key:value`.

```js
assert(!__helpers.removeJSComments(code).match(/:/g))
```

# --seed--

## --seed-contents--

```js
const createPerson = (name, age, gender) => {
  // Only change code below this line
  return {
    name: name,
    age: age,
    gender: gender
  };
  // Only change code above this line
};
```

# --solutions--

```js
const createPerson = (name, age, gender) => {
  // Only change code below this line
  /*return {
    name: name,
    age: age,
    gender: gender
  };*/
  return {
    name,
    age,
    gender
  };
  // Only change code above this line
};
```
