---
id: 587d7b89367417b2b2512b49
title: استخدام تعيين بالتفكيك لتعيين المتغيرات من الكائنات
challengeType: 1
forumTopicId: 301215
dashedName: use-destructuring-assignment-to-assign-variables-from-objects
---

# --description--

يسمح التفكيك بتعيين اسم متغير جديد عند استخراج القيم. يمكنك فعل ذلك بوضع الاسم الجديد بعد النقطتين عند تعيين القيمة.

باستخدام نفس الكائن من المثال السابق:

```js
const user = { name: 'John Doe', age: 34 };
```

إليك كيف يمكنك إعطاء أسماء متغيرات جديدة في التعيين:

```js
const { name: userName, age: userAge } = user;
```

يمكنك قراءتها كـ "احصل على قيمة `user.name` وعيّنها لمتغير جديد اسمه `userName`" وهكذا. ستكون قيمة `userName` هي السلسلة النصية `John Doe`، وقيمة `userAge` هي الرقم `34`.

# --instructions--

استبدل التعيينين بتعيين بالتفكيك مكافئ. يجب أن يظل يعين المتغيرين `highToday` و `highTomorrow` القيمتين `today` و `tomorrow` من كائن `HIGH_TEMPERATURES`.

# --hints--

يجب إزالة صيغة التعيين في ES5.

```js
assert(
  !__helpers.removeJSComments(code).match(/highToday = HIGH_TEMPERATURES\.today/g) &&
    !__helpers.removeJSComments(code).match(/highTomorrow = HIGH_TEMPERATURES\.tomorrow/g)
);
```

يجب استخدام التفكيك لإنشاء المتغير `highToday`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*(today\s*:\s*highToday[^}]*|[^,]*,\s*today\s*:\s*highToday\s*)}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g
  )
);
```

يجب استخدام التفكيك لإنشاء المتغير `highTomorrow`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*(tomorrow\s*:\s*highTomorrow[^}]*|[^,]*,\s*tomorrow\s*:\s*highTomorrow\s*)}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g
  )
);
```

يجب أن تكون قيمة `highToday` مساوية لـ `77` وقيمة `highTomorrow` مساوية لـ `80`.

```js
assert(highToday === 77 && highTomorrow === 80);
```

# --seed--

## --seed-contents--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

// Only change code below this line
  
const highToday = HIGH_TEMPERATURES.today;
const highTomorrow = HIGH_TEMPERATURES.tomorrow; 

// Only change code above this line
```

# --solutions--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

const { today: highToday, tomorrow: highTomorrow } = HIGH_TEMPERATURES;
```
