---
id: 587d7b88367417b2b2512b46
title: تعيين المعلمات الافتراضية لدوالك
challengeType: 1
forumTopicId: 301209
dashedName: set-default-parameters-for-your-functions
---

# --description--

لمساعدتنا في إنشاء دوال أكثر مرونة، يقدم ES6 <dfn>المعلمات الافتراضية</dfn> للدوال.

اطلع على هذه الشفرة:

```js
const greeting = (name = "Anonymous") => "Hello " + name;

console.log(greeting("John"));
console.log(greeting());
```

ستعرض وحدة التحكم السلاسل النصية `Hello John` و `Hello Anonymous`.

تُفعّل المعلمة الافتراضية عندما لا تُحدد المعلمة (أي تكون undefined). كما ترى في المثال أعلاه، ستتلقى المعلمة `name` قيمتها الافتراضية `Anonymous` عندما لا تقدم قيمة للمعلمة. يمكنك إضافة قيم افتراضية لأي عدد من المعلمات تريد.

# --instructions--

عدّل الدالة `increment` بإضافة معلمات افتراضية بحيث تضيف 1 إلى `number` إذا لم تُحدد `value`.

# --hints--

يجب أن تكون نتيجة `increment(5, 2)` هي `7`.

```js
assert(increment(5, 2) === 7);
```

يجب أن تكون نتيجة `increment(5)` هي `6`.

```js
assert(increment(5) === 6);
```

يجب استخدام القيمة الافتراضية `1` للمعلمة `value`.

```js
assert(__helpers.removeJSComments(code).match(/value\s*=\s*1/g));
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const increment = (number, value) => number + value;
// Only change code above this line
```

# --solutions--

```js
const increment = (number, value = 1) => number + value;
```
