---
id: 587d7b89367417b2b2512b48
title: استخدام معامل النشر لتقييم المصفوفات في مكانها
challengeType: 1
forumTopicId: 301222
dashedName: use-the-spread-operator-to-evaluate-arrays-in-place
---

# --description--

يُقدّم ES6 <dfn>معامل النشر</dfn>، الذي يسمح لنا بتوسيع المصفوفات والتعابير الأخرى في الأماكن التي يُتوقع فيها وجود معلمات أو عناصر متعددة.

يستخدم شفرة ES5 أدناه `apply()` لحساب القيمة العظمى في مصفوفة:

```js
var arr = [6, 89, 3, 45];
var maximus = Math.max.apply(null, arr);
```

سيكون لـ `maximus` قيمة `89`.

اضطررنا لاستخدام `Math.max.apply(null, arr)` لأن `Math.max(arr)` تُرجع `NaN`. تتوقع `Math.max()` معلمات مفصولة بفواصل، وليس مصفوفة. يجعل معامل النشر هذه الصيغة أسهل كثيرًا للقراءة والصيانة.

```js
const arr = [6, 89, 3, 45];
const maximus = Math.max(...arr);
```

سيكون لـ `maximus` قيمة `89`.

يرجع `...arr` مصفوفة غير معبأة. بعبارة أخرى، ينشر المصفوفة. مع ذلك، يعمل معامل النشر فقط في المكان، مثل معلمة لدالة أو في مصفوفة حرفية. على سبيل المثال:

```js
const spreaded = [...arr];
```

مع ذلك، الشفرة التالية لن تعمل:

```js
const spreaded = ...arr;
```

# --instructions--

انسخ كل محتويات `arr1` إلى مصفوفة أخرى `arr2` باستخدام معامل النشر.

# --hints--

يجب أن تكون `arr2` نسخة صحيحة من `arr1`.

```js
assert(arr2.every((v, i) => v === arr1[i]) && arr2.length);
```

يجب استخدام معامل النشر `...` لتكرار `arr1`.

```js
assert(__helpers.removeJSComments(code).match(/Array\(\s*\.\.\.arr1\s*\)|\[\s*\.\.\.arr1\s*\]/));
```

يجب أن تبقى `arr2` دون تغيير عندما تتغير `arr1`.

```js
assert((arr1, arr2) => {
  arr1.push('JUN');
  return arr2.length < arr1.length;
});
```

# --seed--

## --seed-contents--

```js
const arr1 = ['JAN', 'FEB', 'MAR', 'APR', 'MAY'];
let arr2;

arr2 = [];  // Change this line

console.log(arr2);
```

# --solutions--

```js
const arr1 = ['JAN', 'FEB', 'MAR', 'APR', 'MAY'];
let arr2;

arr2 = [...arr1];
```
