---
id: 587d7b8d367417b2b2512b59
title: استيراد تصدير افتراضي
challengeType: 1
forumTopicId: 301205
dashedName: import-a-default-export
---

# --description--

في التحدٍّ السابق، تعلّمت عن `export default` واستخداماته. لاستيراد تصدير افتراضي، تحتاج إلى استخدام صيغة `import` مختلفة. في المثال التالي، `add` هو التصدير الافتراضي من ملف `math_functions.js`. إليك كيفية استيراده:

```js
import add from "./math_functions.js";
```

تختلف الصيغة في نقطة رئيسية واحدة. القيمة المستوردة، `add`، ليست محاطة بأقواس متعرجة (`{}`). `add` هنا هو مجرد اسم متغير لأي شيء هو التصدير الافتراضي من ملف `math_functions.js`. يمكنك استخدام أي اسم هنا عند استيراد تصدير افتراضي.

# --instructions--

في الشفرة التالية، استورد التصدير الافتراضي من ملف `math_functions.js` الموجود في نفس الدليل مع هذا الملف. امنح الاستيراد الاسم `subtract`.

# --hints--

يجب أن تستورد `subtract` بشكل صحيح من `math_functions.js`.

```js
assert(__helpers.removeJSComments(code).match(/import\s+subtract\s+from\s+('|")\.\/math_functions\.js\1/g));
```

# --seed--

## --seed-contents--

```js
  
// Only change code above this line

subtract(7,4);
```

# --solutions--

```js
import subtract from "./math_functions.js";

subtract(7,4);
```
