---
id: 587d7b87367417b2b2512b42
title: تعديل مصفوفة معلنة باستخدام const
challengeType: 1
forumTopicId: 301206
dashedName: mutate-an-array-declared-with-const
---

# --description--

إذا لم تكن معتادًا على `const`، اطلع على <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/declare-a-read-only-variable-with-the-const-keyword" target="_blank" rel="noopener noreferrer nofollow">هذا التحدٍّ حول الكلمة المفتاحية <code>const</code></a>.

لإعلان `const` استخدامات كثيرة في جافاسكريبت الحديثة.

يفضل بعض المطوّرين تعيين جميع متغيراتهم باستخدام `const` افتراضيًا، إلا إذا كانوا يعلمون أنهم سيحتاجون لإعادة تعيين القيمة. في هذه الحالة فقط يستخدمون `let`.

مع ذلك، من المهم أن تفهم أن الكائنات (بما في ذلك المصفوفات والدوال) المعينة لمتغير باستخدام `const` لا تزال قابلة للتغيير. إعلان `const` يمنع فقط إعادة تعيين معرف المتغير.

```js
const s = [5, 6, 7];
s = [1, 2, 3];
s[2] = 45;
console.log(s);
```

سيسبب `s = [1, 2, 3]` خطأً. بعد تعليق هذا السطر، ستعرض وحدة التحكم `console.log` القيمة `[5, 6, 45]`.

كما ترى، يمكنك تعديل الكائن `[5, 6, 7]` نفسه، وسيظل المتغير `s` يشير إلى المصفوفة المعدلة `[5, 6, 45]`. مثل جميع المصفوفات، عناصر المصفوفة في `s` قابلة للتغيير، لكن لأن `const` استُخدم، لا يمكنك استخدام معرف المتغير `s` للإشارة إلى مصفوفة مختلفة باستخدام معامل التعيين.

# --instructions--

تم إعلان مصفوفة كـ `const s = [5, 7, 2]`. غيّر المصفوفة إلى `[2, 5, 7]` باستخدام تعيينات لعناصر مختلفة.

# --hints--

يجب ألا تستبدل الكلمة المفتاحية `const`.

```js
assert(__helpers.removeJSComments(code).match(/const/g));
```

يجب أن يكون `s` متغيرًا ثابتًا (باستخدام `const`).

```js
assert(__helpers.removeJSComments(code).match(/const\s+s/g));
```

يجب ألا تغيّر إعلان المصفوفة الأصلي.

```js
assert(__helpers.removeJSComments(code).match(
/const\s+s\s*=\s*\[\s*5\s*,\s*7\s*,\s*2\s*\]\s*;?/g
));
```

يجب أن تكون قيمة `s` مساوية لـ `[2, 5, 7]`.

```js
assert.deepEqual(s, [2, 5, 7]);
```

# --seed--

## --seed-contents--

```js
const s = [5, 7, 2];
function editInPlace() {
  // Only change code below this line

  // Using s = [2, 5, 7] would be invalid

  // Only change code above this line
}
editInPlace();
```

# --solutions--

```js
const s = [5, 7, 2];
function editInPlace() {
  s[0] = 2;
  s[1] = 5;
  s[2] = 7;
}
editInPlace();
```
