---
id: 598f48a36c8c40764b4e52b3
title: منع تحوُّل الكائن
challengeType: 1
forumTopicId: 301207
dashedName: prevent-object-mutation
---

# --description--

كما رأيت في التحدٍّ السابق، إعلان `const` وحده لا يحمي بياناتك من التحوُّل. لضمان عدم تغيير بياناتك، توفر JavaScript دالة `Object.freeze` لمنع تحوُّل البيانات.

أي محاولة لتغيير الكائن ستُرفض، مع رمي خطأ إذا كان السكربت يعمل في الوضع الصارم.

```js
let obj = {
  name:"FreeCodeCamp",
  review:"Awesome"
};
Object.freeze(obj);
obj.review = "bad";
obj.newProp = "Test";
console.log(obj); 
```

تعيينات `obj.review` و `obj.newProp` ستؤدي إلى أخطاء، لأن محررنا يعمل في الوضع الصارم افتراضيًا، وستعرض وحدة التحكم القيمة `{ name: "FreeCodeCamp", review: "Awesome" }`.

# --instructions--

في هذا التحدٍّ ستستخدم `Object.freeze` لمنع ثوابت الرياضيات من التغيير. يجب تجميد كائن `MATH_CONSTANTS` حتى لا يتمكن أحد من تعديل قيمة `PI` أو إضافة أو حذف خصائص.

# --hints--

لا يجب استبدال كلمة `const`.

```js
assert(__helpers.removeJSComments(code).match(/const/g));
```

يجب أن يكون `MATH_CONSTANTS` متغيرًا ثابتًا (باستخدام `const`).

```js
assert(__helpers.removeJSComments(code).match(/const\s+MATH_CONSTANTS/g));
```

لا يجب تغيير الإعلان الأصلي لـ `MATH_CONSTANTS`.

```js
assert(__helpers.removeJSComments(code).match(
   /const\s+MATH_CONSTANTS\s+=\s+{\s+PI:\s+3.14\s+};/g
));
```

يجب أن تساوي `PI` القيمة `3.14`.

```js
assert(PI === 3.14);
```

# --seed--

## --seed-contents--

```js
function freezeObj() {
  const MATH_CONSTANTS = {
    PI: 3.14
  };
  // Only change code below this line


  // Only change code above this line
  try {
    MATH_CONSTANTS.PI = 99;
  } catch(ex) {
    console.log(ex);
  }
  return MATH_CONSTANTS.PI;
}
const PI = freezeObj();
```

# --solutions--

```js
function freezeObj() {
  const MATH_CONSTANTS = {
    PI: 3.14
  };
  Object.freeze(MATH_CONSTANTS);

  try {
    MATH_CONSTANTS.PI = 99;
  } catch(ex) {
    console.log(ex);
  }
  return MATH_CONSTANTS.PI;
}
const PI = freezeObj();
```
