---
id: 5cdafbc32913098997531680
title: إكمال وعد باستخدام resolve و reject
challengeType: 1
forumTopicId: 301196
dashedName: complete-a-promise-with-resolve-and-reject
---

# --description--

للوعد ثلاث حالات: `pending`، و`fulfilled`، و`rejected`. الوعد الذي أنشأته في التحدٍّ السابق يبقى عالقًا في حالة `pending` إلى الأبد لأنك لم تضف طريقة لإكمال الوعد. تُستخدم المعلمتان `resolve` و`reject` المعطاة للمعلمة الخاصة بالوعد لتحقيق ذلك. تُستخدم `resolve` عندما تريد أن ينجح الوعد، وتُستخدم `reject` عندما تريد فشله. هاتان طريقتان تأخذان معلمة، كما هو موضح أدناه.

```js
const myPromise = new Promise((resolve, reject) => {
  if(condition here) {
    resolve("Promise was fulfilled");
  } else {
    reject("Promise was rejected");
  }
});
```

المثال أعلاه يستخدم سلاسل نصية كمعلمة لهذه الدوال، لكنها يمكن أن تكون أي شيء فعليًا. غالبًا ما تكون كائنًا تستخدم بياناته لوضعها على موقعك أو في مكان آخر.

# --instructions--

اجعل الوعد يتعامل مع النجاح والفشل. إذا كانت قيمة `responseFromServer` هي `true`، فاستدعِ طريقة `resolve` لإكمال الوعد بنجاح. مرر إلى `resolve` سلسلة نصية بقيمة `We got the data`. إذا كانت قيمة `responseFromServer` هي `false`، فاستخدم طريقة `reject` بدلاً من ذلك ومرر إليها السلسلة النصية: `Data not received`.

# --hints--

يجب استدعاء `resolve` مع السلسلة النصية المتوقعة عندما تكون حالة `if` هي `true`.

```js
assert(
  __helpers.removeJSComments(code).match(/if\s*\(\s*responseFromServer\s*\)\s*{\s*resolve\s*\(\s*('|"|`)We got the data\1\s*\)(\s*|\s*;\s*)}/g)
);
```

يجب استدعاء `reject` مع السلسلة النصية المتوقعة عندما تكون حالة `if` هي `false`.

```js
assert(
  __helpers.removeJSComments(code).match(/}\s*else\s*{\s*reject\s*\(\s*('|"|`)Data not received\1\s*\)(\s*|\s*;\s*)}/g)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer represents a response from a server
  let responseFromServer;
    
  if(responseFromServer) {
    // Change this line
  } else {  
    // Change this line
  }
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer represents a response from a server
  let responseFromServer;

  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});
```
