---
id: 5cdafbe72913098997531682
title: التعامل مع وعد مرفوض باستخدام catch
challengeType: 1
forumTopicId: 301204
dashedName: handle-a-rejected-promise-with-catch
---

# --description--

`catch` هي الطريقة المستخدمة عندما يُرفض الوعد الخاص بك. تُنفذ مباشرة بعد استدعاء طريقة `reject` الخاصة بالوعد. هذه هي الصيغة النحوية:

```js
myPromise.catch(error => {
  
});
```

`error` هي المعلمة التي تُمرر إلى طريقة `reject`.

# --instructions--

أضف طريقة `catch` إلى وعدك. استخدم `error` كمعلمة في دالة الاستدعاء الراجع الخاصة بها وسجل `error` في وحدة التحكم.

# --hints--

يجب أن تستدعي طريقة `catch` على الوعد.

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/(makeServerRequest|\))\.catch\(/g)
);
```

يجب أن تحتوي طريقة `catch` على دالة استدعاء راجع بمعلمة `error`.

```js
assert(
  /\.catch\((function\(error\){|error|\(error\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

يجب أن تسجل `error` في وحدة التحكم.

```js
assert(
  /\.catch\((function\(error\){|error|\(error\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  ) &&
    __helpers
      .removeWhiteSpace(__helpers.removeJSComments(code))
      .match(/\.catch\(.*?error.*?console.log\(error\).*?\)/)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to false to represent an unsuccessful response from a server
  let responseFromServer = false;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to false to represent an unsuccessful response from a server
  let responseFromServer = false;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});

makeServerRequest.catch(error => {
  console.log(error);
});
```
