---
id: 5cdafbd72913098997531681
title: التعامل مع وعد مُنجز باستخدام then
challengeType: 1
forumTopicId: 301203
dashedName: handle-a-fulfilled-promise-with-then
---

# --description--

تُفيد الوعود أكثر عندما يكون لديك عملية تستغرق وقتًا غير معروف في شفرتك (أي شيء غير متزامن)، وغالبًا ما يكون طلب خادم. عندما تُجري طلبًا إلى خادم، يستغرق بعض الوقت، وبعد الانتهاء عادةً تريد فعل شيء ما مع الاستجابة من الخادم. يمكن تحقيق ذلك باستخدام طريقة `then`.

```js
Promise.prototype.then(onFulfilled, onRejected)
```

تُجدول طريقة `then` دوال استدعاء راجع لإكمال الوعد في المستقبل - إما تحقيقه أو رفضه. سيتم تنفيذ أحد معالجي `onFulfilled` أو `onRejected` للتعامل مع تحقيق الوعد الحالي أو رفضه. عندما يُحقق الوعد باستخدام `resolve`، يُستدعى معالج `onFulfilled`.

```js
myPromise.then(result => {
  
});
```

يأتي `result` من المعلمة المعطاة إلى طريقة `resolve`.

# --instructions--

أضف طريقة `then` إلى وعدك. استخدم `result` كمعلمة لدالة الاستدعاء الراجع الخاصة بها وسجل `result` في وحدة التحكم.

# --hints--

يجب أن تستدعي طريقة `then` على الوعد.

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/(makeServerRequest|\))\.then\(/g)
);
```

يجب أن تحتوي طريقة `then` على دالة استدعاء راجع بمعلمة `result`.

```js
assert(
  /\.then\((function\(result\){|result|\(result\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

يجب أن تسجل `result` في وحدة التحكم.

```js
assert(
  /\.then\((function\(result\){|result|\(result\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  ) &&
    __helpers
      .removeWhiteSpace(__helpers.removeJSComments(code))
      .match(/\.then\(.*?result.*?console.log\(result\).*?\)/)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to true to represent a successful response from a server
  let responseFromServer = true;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to true to represent a successful response from a server
  let responseFromServer = true;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});
```
