---
id: 5cdafbb0291309899753167f
title: إنشاء وعد في JavaScript
challengeType: 1
forumTopicId: 301197
dashedName: create-a-javascript-promise
---

# --description--

الوعد في JavaScript هو بالضبط ما يوحي به اسمه - تستخدمه لتقديم وعد بفعل شيء ما، عادةً بشكل غير متزامن. عند اكتمال المهمة، إما أن تُفي بوعدك أو تفشل في ذلك. `Promise` هي دالة بناء، لذا تحتاج إلى استخدام الكلمة المفتاحية `new` لإنشائها. تأخذ دالة كمعامل لها، مع معلمتين - `resolve` و `reject`. هاتان طريقتان تُستخدمان لتحديد نتيجة الوعد. الصيغة النحوية تبدو هكذا:

```js
const myPromise = new Promise((resolve, reject) => {

});
```

# --instructions--

أنشئ وعدًا جديدًا باسم `makeServerRequest`. مرر دالة تحتوي على معلمتي `resolve` و `reject` إلى دالة البناء.

# --hints--

يجب تعيين وعد إلى متغير معلن اسمه `makeServerRequest`.

```js
assert(makeServerRequest instanceof Promise);
```

يجب أن يستقبل وعدك دالة تحتوي على معلمتي `resolve` و `reject`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /Promise\s*\(\s*(function\s*\(\s*resolve\s*,\s*reject\s*\)\s*{|\(\s*resolve\s*,\s*reject\s*\)\s*=>\s*{)[^}]*}/g
  )
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {

});
```
