---
id: 5cddbfd622f1a59093ec611d
title: إنشاء سكربت وحدة
challengeType: 6
forumTopicId: 301198
dashedName: create-a-module-script
---

# --description--

بدأ JavaScript بدور صغير في موقع ويب يعتمد في الغالب على HTML. اليوم، أصبح كبيرًا، وبعض المواقع تُبنى بالكامل تقريبًا باستخدام JavaScript. لجعل JavaScript أكثر تنظيمًا ونظافة وقابلية للصيانة؛ قدم ES6 طريقة سهلة لمشاركة الكود بين ملفات JavaScript. يتضمن ذلك تصدير أجزاء من ملف لاستخدامها في ملف أو أكثر، واستيراد الأجزاء التي تحتاجها، حيث تحتاجها. للاستفادة من هذه الميزة، عليك إنشاء سكربت في مستند HTML الخاص بك مع `type` بقيمة `module`. إليك مثالًا:

```html
<script type="module" src="filename.js"></script>
```

يمكن للسكربت الذي يستخدم نوع `module` الآن استخدام ميزتي `import` و `export` التي ستتعلمها في التحديات القادمة.

# --instructions--

أضف سكربتًا إلى مستند HTML من نوع `module` وحدد له ملف المصدر `index.js`

# --hints--

يجب إنشاء وسم `script`.

```js
assert(__helpers.removeJSComments(code).match(/<\s*script[^>]*>\s*<\/\s*script\s*>/g));
```

يجب أن يحتوي وسم `script` على الخاصية `type` بقيمة `module`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /<\s*script\s+[^t]*type\s*=\s*('|")module\1[^>]*>\s*<\/\s*script\s*>/g
  )
);
```

يجب أن يحتوي وسم `script` على `src` بقيمة `index.js`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /<\s*script\s+[^s]*src\s*=\s*('|")index\.js\1[^>]*>\s*<\/\s*script\s*>/g
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <!-- Only change code below this line -->

    <!-- Only change code above this line -->
  </body>
</html>
```

# --solutions--

```html
<html>
  <body>
    <script type="module" src="index.js"></script>
  </body>
</html>
```
