---
id: 5cfa550e84205a357704ccb6
title: استخدام تعيين بالتفكيك لاستخراج القيم من الكائنات
challengeType: 1
forumTopicId: 301216
dashedName: use-destructuring-assignment-to-extract-values-from-objects
---

# --description--

<dfn>تعيين بالتفكيك</dfn> هو صيغة نحوية خاصة أُدخلت في ES6، لتعيين القيم المأخوذة مباشرة من كائن بطريقة مرتبة.

انظر إلى شفرة ES5 التالية:

```js
const user = { name: 'John Doe', age: 34 };

const name = user.name;
const age = user.age;
```

سيكون لـ `name` قيمة السلسلة النصية `John Doe`، و`age` سيكون له الرقم `34`.

إليك بيان تعيين مكافئ باستخدام صيغة التفكيك في ES6:

```js
const { name, age } = user;
```

مرة أخرى، سيكون لـ `name` قيمة السلسلة النصية `John Doe`، و`age` سيكون له الرقم `34`.

هنا، ستُنشأ المتغيرات `name` و`age` وتُعيَّن القيم الخاصة بها من كائن `user`. يمكنك أن ترى مدى نظافة هذه الطريقة.

يمكنك استخراج عدد كبير أو قليل من القيم من الكائن كما تريد.

# --instructions--

استبدل التعيينَين بتعيين بالتفكيك مكافئ. يجب أن يظل يعين المتغيرين `today` و`tomorrow` القيم الخاصة بـ `today` و`tomorrow` من كائن `HIGH_TEMPERATURES`.

# --hints--

يجب إزالة صيغة التعيين في ES5.

```js
assert(
  !__helpers.removeJSComments(code).match(/today\s*=\s*HIGH_TEMPERATURES\.(today|tomorrow)/g)
);
```

يجب استخدام التفكيك لإنشاء المتغير `today`.

```js
assert(
  __helpers.removeJSComments(code).match(/(var|let|const)\s*{\s*(today[^}]*|[^,]*,\s*today)\s*}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g)
);
```

يجب استخدام التفكيك لإنشاء المتغير `tomorrow`.

```js
assert(
  __helpers.removeJSComments(code).match(/(var|let|const)\s*{\s*(tomorrow[^}]*|[^,]*,\s*tomorrow)\s*}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g)
);
```

يجب أن تكون قيمة `today` مساوية لـ `77` وقيمة `tomorrow` مساوية لـ `80`.

```js
assert(today === 77 && tomorrow === 80);
```

# --seed--

## --seed-contents--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

// Only change code below this line

const today = HIGH_TEMPERATURES.today;
const tomorrow = HIGH_TEMPERATURES.tomorrow;

// Only change code above this line
```

# --solutions--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

const { today, tomorrow } = HIGH_TEMPERATURES;
```
