---
id: 587d78ad367417b2b2512af8
title: Alinhar elementos usando a propriedade align-items
challengeType: 0
forumTopicId: 301101
dashedName: align-elements-using-the-align-items-property
---

# --description--

A propriedade `align-items` é similar à `justify-content`. Lembre que a propriedade `justify-content` alinhou os flex items ao longo do eixo principal. Para linhas, o eixo principal é uma linha horizontal e para colunas é uma linha vertical.

Contêineres flex também possuem um **eixo cruzado**, que é o oposto do eixo principal. Para linhas, o eixo transversal é vertical, e para colunas, o eixo transversal é horizontal.

O CSS oferece a propriedade `align-items` para alinhar os flex items ao longo do eixo transversal. Para linhas, ela indica ao CSS como empurrar os itens para cima ou para baixo dentro do container. E para uma coluna, como empurrar todos os itens para a esquerda ou direita dentro do container.

Os valores disponíveis para `align-items` são:

<ul><li><code>flex-start</code>: alinha os itens ao início do container flex. Para linhas, isso alinha os itens ao topo do container. Para colunas, isso alinha os itens à esquerda do container.</li><li><code>flex-end</code>: alinha os itens ao final do container flex. Para linhas, isso alinha os itens à parte inferior do container. Para colunas, isso alinha os itens à direita do container.</li><li><code>center</code>: alinha os itens ao centro. Para linhas, isso alinha verticalmente os itens (espaço igual acima e abaixo dos itens). Para colunas, isso alinha horizontalmente (espaço igual à esquerda e à direita dos itens).</li><li><code>stretch</code>: estica os itens para preencher o container flex. Por exemplo, itens em linhas são esticados para preencher o container flex de cima para baixo. Esse é o valor padrão se nenhum valor de <code>align-items</code> for especificado.</li><li><code>baseline</code>: alinha os itens às suas linhas de base. Linha de base é um conceito de texto, pense nela como a linha sobre a qual as letras se apoiam.</li></ul>

# --instructions--

Alguns exemplos vão ajudar você a entender os valores dessa propriedade em ação. Adicione a propriedade CSS `align-items` ao elemento `#box-container` e dê a ela o valor de `center`.

**Bônus**Experimente as outras opções para a propriedade `align-items` no editor de código para ver suas diferenças. Mas observe que o valor `center` será o único que vai fazer você passar neste desafio.

# --hints--

O elemento de id `#box-container` deve ter a propriedade `align-items` definida com o valor de `center`.

```js
const boxContainer = document.querySelector('#box-container');
const alignment = window.getComputedStyle(boxContainer)['align-items'];
assert.strictEqual(alignment, 'center');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 200px;
    font-size: 24px;
  }

  #box-2 {
    background-color: orangered;
    width: 200px;
    font-size: 18px;
  }
</style>

<div id="box-container">
  <div id="box-1"><p>Hello</p></div>
  <div id="box-2"><p>Goodbye</p></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;
    align-items: center;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 200px;
    font-size: 24px;
  }

  #box-2 {
    background-color: orangered;
    width: 200px;
    font-size: 18px;
  }
</style>

<div id="box-container">
  <div id="box-1"><p>Hello</p></div>
  <div id="box-2"><p>Goodbye</p></div>
</div>
```
