---
id: a6b0bb188d873cb2c8729495
title: Tekeleza mabadiliko ya herufi maalum kuwa entiti za HTML
challengeType: 26
dashedName: implement-an-html-entity-converter
---

# --description--

Maabara hii ni kuhusu kubadilisha herufi maalum katika mfuatano wa herufi kuwa entiti zao za HTML zinazolingana.

**Lengo**: Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara hii.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kitendakazi cha `convertHTML` kinachokubali mfuatano wa herufi kama hoja.
1. Kitendakazi cha `convertHTML` kinapaswa kurudisha mfuatano mpya wa herufi kwa kubadilisha herufi maalum katika mfuatano wa hoja kuwa entiti zao za HTML zinazolingana.

   - `&` inapaswa kubadilishwa kuwa `&amp;`.
   - `<` inapaswa kubadilishwa kuwa `&lt;`.
   - `>` inapaswa kubadilishwa kuwa `&gt;`.
   - `"` inapaswa kubadilishwa kuwa `&quot;`.
   - `'` inapaswa kubadilishwa kuwa `&apos;`.

# --hints--

Unapaswa kuwa na kitendakazi cha `convertHTML`.

```js
assert.isFunction(convertHTML);
```

`convertHTML("Dolce & Gabbana")` inapaswa kurudisha mfuatano wa herufi `Dolce &amp; Gabbana`.

```js
assert.match(convertHTML('Dolce & Gabbana'), /Dolce &amp; Gabbana/);
```

`convertHTML("Hamburgers < Pizza < Tacos")` inapaswa kurudisha mfuatano wa herufi `Hamburgers &lt; Pizza &lt; Tacos`.

```js
assert.match(
  convertHTML('Hamburgers < Pizza < Tacos'),
  /Hamburgers &lt; Pizza &lt; Tacos/
);
```

`convertHTML("Sixty > twelve")` inapaswa kurudisha mfuatano wa herufi `Sixty &gt; twelve`.

```js
assert.match(convertHTML('Sixty > twelve'), /Sixty &gt; twelve/);
```

`convertHTML('Stuff in "quotation marks"')` inapaswa kurudisha mfuatano wa herufi `Stuff in &quot;quotation marks&quot;`.

```js
assert.match(
  convertHTML('Stuff in "quotation marks"'),
  /Stuff in &quot;quotation marks&quot;/
);
```

`convertHTML("Schindler's List")` inapaswa kurudisha mfuatano wa herufi `Schindler&apos;s List`.

```js
assert.match(convertHTML("Schindler's List"), /Schindler&apos;s List/);
```

`convertHTML("<>")` inapaswa kurudisha mfuatano wa herufi `&lt;&gt;`.

```js
assert.match(convertHTML('<>'), /&lt;&gt;/);
```

`convertHTML("abc")` inapaswa kurudisha mfuatano wa herufi `abc`.

```js
assert.strictEqual(convertHTML('abc'), 'abc');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const map = {
  '&': '&amp;',
  '<': '&lt;',
  '>': '&gt;',
  '"': '&quot;',
  "'": '&apos;'
};

function convertHTML(str) {
  return str.replace(/[&<>"']/g, function(char) {
    return map[char];
  });
}
```
