---
id: 68a47dcdbce12193caf77012
title: Implementa un Buscador de Índices Ordenados
challengeType: 26
dashedName: lab-sorted-index-finder
---

# --description--

En este laboratorio crearás una función que devuelve el índice más bajo en el que un valor debe ser insertado en un arreglo una vez que ha sido ordenado en orden ascendente.

**Objetivo**: Cumple las historias de usuario a continuación y logra que todas las pruebas pasen para completar el laboratorio.

**Historias de usuario:**

1. Debes tener una función `getIndexToIns` que tome dos argumentos: un arreglo y un número.
2. Debes usar el método `sort` para ordenar el arreglo en orden ascendente.
3. Tu función `getIndexToIns` debería devolver el índice más bajo en el que se debe insertar el número usando el método `findIndex`.
4. Tu función `getIndexToIns` siempre debe devolver un número.

**Hint:**  
The `findIndex` method is a built-in array method in JavaScript. It takes a callback function and returns the index of the first element that satisfies the condition. Both `findIndex` and `sort` are higher-order functions.

**Ejemplos:**

- `getIndexToIns([1, 2, 3, 4], 1.5)` debería devolver `1` porque `1.5` es mayor que `1` (índice `0`) y menor que `2` (índice `1`).
- `getIndexToIns([20, 3, 5], 19)` debería devolver `2` porque después de ordenar a `[3, 5, 20]`, `19` es menor que `20` (índice `2`) y mayor que `5` (índice `1`).


# --hints--

Debes tener una función `getIndexToIns`.

```js
assert.isFunction(getIndexToIns);
```

`getIndexToIns` siempre debe devolver un número.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

Tu función debería usar el método `sort`.

```js
assert.match(getIndexToIns.toString(), /\.sort\(/);
```


Tu función debe usar el método `findIndex`.

```js
assert.match(getIndexToIns.toString(), /\.findIndex\(/);
```

`getIndexToIns([10, 20, 30, 40, 50], 35)` debe devolver `3`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` debe devolver `2`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

`getIndexToIns([40, 60], 50)` debe devolver `1`.

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

`getIndexToIns([3, 10, 5], 3)` debe devolver `0`.

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

`getIndexToIns([5, 3, 20, 3], 5)` debe devolver `2`.

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

`getIndexToIns([2, 20, 10], 19)` debe devolver `2`.

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

`getIndexToIns([3, 10, 5], 11)` debería devolver `3`

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 11), 3);
```

`getIndexToIns([], 5)` debería devolver `0`

```js
assert.strictEqual(getIndexToIns([], 5), 0);
```

# --seed--

## --seed-contents--

```js
``` 

# --solutions--

```js
function getIndexToIns(arr, num) {
  const sorted = [...arr].sort((a, b) => a - b);
  const index = sorted.findIndex(el => el >= num);

  return index === -1 ? sorted.length : index;
}
``` 

