---
id: 5e44414f903586ffb414c950
title: Ein Projekt zum Erstellen eines Wahrscheinlichkeitsrechners
challengeType: 23
saveSubmissionToDB: true
forumTopicId: 462364
dashedName: build-a-probability-calculator-project
---

# --description--

Angenommen, es gibt einen Hut mit 5 blauen Kugeln, 4 roten Kugeln und 2 grünen Kugeln. Wie hoch ist die Wahrscheinlichkeit, dass bei einer zufälligen Ziehung von 4 Kugeln mindestens 1 rote Kugel und 2 grüne Kugeln enthalten sind? Obwohl es möglich wäre, die Wahrscheinlichkeit mit fortgeschrittener Mathematik zu berechnen, ist eine einfachere Methode, ein Programm zu schreiben, das eine große Anzahl von Experimenten durchführt, um eine ungefähre Wahrscheinlichkeit zu schätzen.

Für dieses Projekt schreiben Sie ein Programm, das die ungefähre Wahrscheinlichkeit bestimmt, bestimmte Kugeln zufällig aus einem Hut zu ziehen.

Erstellen Sie zunächst eine `Hat`-Klasse in `main.py`. Die Klasse sollte eine variable Anzahl von Argumenten akzeptieren, die die Anzahl der Kugeln jeder Farbe im Hut angeben. Zum Beispiel könnte ein Klassenobjekt auf eine der folgenden Arten erstellt werden:

```py
hat1 = Hat(yellow=3, blue=2, green=6)
hat2 = Hat(red=5, orange=4)
hat3 = Hat(red=5, orange=4, black=1, blue=0, pink=2, striped=9)
```

Ein Hut wird immer mit mindestens einer Kugel erstellt. Die beim Erstellen an das Hut-Objekt übergebenen Argumente sollten in eine `contents`-Instanzvariable umgewandelt werden. `contents` sollte eine Liste von Strings sein, die für jede Kugel im Hut einen Eintrag enthält. Jeder Eintrag in der Liste sollte einen Farbnamen darstellen, der eine einzelne Kugel dieser Farbe repräsentiert. Wenn Ihr Hut zum Beispiel `{'red': 2, 'blue': 1}` ist, sollte `contents` `['red', 'red', 'blue']` sein.

Die `Hat`-Klasse sollte eine `draw`-Methode haben, die ein Argument akzeptiert, das angibt, wie viele Kugeln aus dem Hut gezogen werden sollen. Diese Methode sollte Kugeln zufällig aus `contents` entfernen und diese Kugeln als Liste von Strings zurückgeben. Die Kugeln sollten während der Ziehung nicht zurück in den Hut gelegt werden, ähnlich einem Urnenexperiment ohne Zurücklegen. Wenn die Anzahl der zu ziehenden Kugeln die verfügbare Menge übersteigt, geben Sie alle Kugeln zurück.

Erstellen Sie anschließend eine `experiment`-Funktion in `main.py` (nicht innerhalb der `Hat`-Klasse). Diese Funktion sollte die folgenden Argumente akzeptieren:

- `hat`: Ein Hut-Objekt, das innerhalb der Funktion kopiert werden sollte.
- `expected_balls`: Ein Objekt, das die genaue Gruppe von Kugeln angibt, die im Experiment aus dem Hut gezogen werden sollen. Zum Beispiel, um die Wahrscheinlichkeit zu bestimmen, 2 blaue Kugeln und 1 rote Kugel aus dem Hut zu ziehen, setzen Sie `expected_balls` auf `{'blue':2, 'red':1}`.
- `num_balls_drawn`: Die Anzahl der Kugeln, die in jedem Experiment aus dem Hut gezogen werden.
- `num_experiments`: Die Anzahl der durchzuführenden Experimente. (Je mehr Experimente durchgeführt werden, desto genauer wird die ungefähre Wahrscheinlichkeit.)

Die `experiment`-Funktion sollte eine Wahrscheinlichkeit zurückgeben.

Zum Beispiel, wenn Sie die Wahrscheinlichkeit bestimmen möchten, mindestens zwei rote Kugeln und eine grüne Kugel zu erhalten, wenn Sie fünf Kugeln aus einem Hut ziehen, der sechs schwarze, vier rote und drei grüne Kugeln enthält. Dazu führen Sie `N` Experimente durch, zählen, wie oft `M` Sie mindestens zwei rote Kugeln und eine grüne Kugel erhalten, und schätzen die Wahrscheinlichkeit als `M/N`. Jedes Experiment besteht darin, mit einem Hut zu beginnen, der die angegebenen Kugeln enthält, mehrere Kugeln zu ziehen und zu prüfen, ob Sie die Kugeln erhalten haben, die Sie zu ziehen versucht haben.

So würden Sie die `experiment`-Funktion basierend auf dem obigen Beispiel mit 2000 Experimenten aufrufen:

```py
hat = Hat(black=6, red=4, green=3)
probability = experiment(hat=hat,
                  expected_balls={'red':2,'green':1},
                  num_balls_drawn=5,
                  num_experiments=2000)
```

Die Ausgabe wäre ungefähr so:

```bash
0.356
```

Da dies auf zufälligen Ziehungen basiert, wird die Wahrscheinlichkeit bei jedem Ausführen des Codes leicht unterschiedlich sein.

_Hinweis: Verwenden Sie die Module, die bereits oben importiert sind. Initialisieren Sie den Zufalls-Seed nicht innerhalb der Datei._

Hinweis: Öffnen Sie die Browser-Konsole mit F12, um eine ausführlichere Ausgabe der Tests zu sehen.

# --hints--

Die Erstellung eines `hat`-Objekts sollte den korrekten Inhalt hinzufügen.

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_class_contents(self):
        hat = probability_calculator.Hat(red=3,blue=2)
        actual = hat.contents
        expected = ["red","red","red","blue","blue"]
        self.assertEqual(actual, expected, 'Expected creation of hat object to add correct contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

Die `draw`-Methode in der `hat`-Klasse sollte die Anzahl der Elemente im Inhalt reduzieren.


```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)

class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_draw(self):
        hat = probability_calculator.Hat(red=5,blue=2)
        actual = hat.draw(2)
        expected = ['blue', 'red']
        self.assertEqual(actual, expected, 'Expected hat draw to return two random items from hat contents.')
        actual = len(hat.contents)
        expected = 5
        self.assertEqual(actual, expected, 'Expected hat draw to reduce number of items in contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

Die `draw`-Methode sollte korrekt funktionieren, wenn die Anzahl der zu entnehmenden Kugeln größer ist als die Anzahl der Kugeln im Hut.

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_draw_2(self):
        hat = probability_calculator.Hat(yellow=5,red=1,green=3,blue=9,test=1)
        actual = sorted(hat.draw(20))
        expected = sorted(['yellow', 'yellow', 'yellow', 'yellow', 'yellow', 'red', 'green', 'green', 'green', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'test'])
        self.assertEqual(actual, expected, 'Expected hat draw to return all items from hat contents.')
        actual = len(hat.contents)
        expected = 0
        self.assertEqual(actual, expected, 'Expected hat draw to leave no items in contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

Die `experiment`-Methode sollte eine unterschiedliche Wahrscheinlichkeit zurückgeben.


```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_prob_experiment(self):
        hat = probability_calculator.Hat(blue=3,red=2,green=6)
        probability = probability_calculator.experiment(hat=hat, expected_balls={"blue":2,"green":1}, num_balls_drawn=4, num_experiments=1000)
        actual = probability
        expected = 0.272
        self.assertAlmostEqual(actual, expected, delta = 0.01, msg = 'Expected experiment method to return a different probability.')
        hat = probability_calculator.Hat(yellow=5,red=1,green=3,blue=9,test=1)
        probability = probability_calculator.experiment(hat=hat, expected_balls={"yellow":2,"blue":3,"test":1}, num_balls_drawn=20, num_experiments=100)
        actual = probability
        expected = 1.0
        self.assertAlmostEqual(actual, expected, delta = 0.01, msg = 'Expected experiment method to return a different probability.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

# --seed--

## --seed-contents--

```py
import copy
import random

class Hat:
    pass

def experiment(hat, expected_balls, num_balls_drawn, num_experiments):
    pass
```

# --solutions--

```py
import copy
import random

class Hat:
    def __init__(self, **hat):
        self.hat = hat
        contents = []
        for i in hat:
            for j in range(hat[i]):
                contents.append(i)           
        self.contents = contents
               

    def draw(self, number): 
        drawn = []
        if number >= len(self.contents):
            drawn.extend(self.contents)
            self.contents = []
        else:
            for i in range(number):
                drawn.append(
                    self.contents.pop(random.randrange(len(self.contents)))
                )                
        return drawn

def experiment(hat, expected_balls, num_balls_drawn, num_experiments):    
    expected_balls_list = []
    drawn_list = []
    success = 0
    for i in expected_balls:
        for j in range(expected_balls[i]):
            expected_balls_list.append(i)
    for j in range(num_experiments):
        hat_copy = copy.deepcopy(hat)
        drawn_list.append(hat_copy.draw(num_balls_drawn))        
        exp_ball_list_copy = expected_balls_list[:]
        for k in range(len(drawn_list[j])):
            try:
                ind = exp_ball_list_copy.index(drawn_list[j][k])
                exp_ball_list_copy.pop(ind)
            except:
                continue
        if len(exp_ball_list_copy) == 0:
            success += 1        
    probability = success/num_experiments    
    return probability
```
