---
id: 60a3e3396c7b40068ad69970
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

Before you can start styling the `div` you added, you need to link your CSS to your HTML.

添加 `link` 元素以链接你的 `styles.css` 文件。 将 `href` 设置为 `styles.css`，并记住将 `rel` 属性设置为 `stylesheet`。

# --hints--

你的代码应该有一个 `link` 元素。

```js
assert(/<link/.test(code))
```

`link` 元素应该有一个 `rel` 属性，其值为 `stylesheet`。

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

`link` 元素应该有一个 `href` 属性，其值为 `styles.css`。

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
