---
id: 699d20623e57e8ffee885455
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Como você pode ver na saída no console, depois de dividir, cada parte ainda tem espaços em branco extras no início e/ou no fim. Para limpar cada elemento, você precisa processar cada um individualmente.

Este é um caso perfeito para um loop `for`, que repete um bloco de código uma vez para cada elemento em um iterável, como um array:

```js
for (let i = 0; i < array.length; i++) {
  // code to run on each iteration
}
```

`i` começa em `0` (o primeiro índice), o loop continua enquanto `i` for menor que o comprimento do array, e `i++` aumenta `i` em um após cada iteração.

Dentro de `parseCard`, declare um array vazio chamado `trimmedParts`. Depois adicione um loop `for` que itere sobre `parts`. Deixe o corpo do loop vazio por enquanto, pois você vai preenchê-lo no próximo passo.

# --hints--

Você deve declarar uma variável chamada `trimmedParts` definida como um array vazio.

```js
assert.match(parseCard.toString(), /(const|let|var)\s+trimmedParts\s*=\s*\[\s*\]/);
```

Você deve usar um loop `for` para iterar sobre `parts`.

```js
assert.match(parseCard.toString(), /for\s*\(\s*(var|let)\s+(\w+)\s*=\s*0\s*;\s*\2\s*<\s*parts\s*\.\s*length\s*;\s*\2\+\+\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
--fcc-editable-region--

--fcc-editable-region--
  return parts;
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
