---
id: 699d20623e57e8ffee885456
title: passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Agora que as partes estão limpas, você pode transformá-las em um objeto estruturado. Em JavaScript, objetos usam chaves com pares "chave: valor". Por exemplo, `{ title: "Dune", year: 1965 }`.

Declare quatro variáveis: `title`, `author`, `year` e `location`. Atribua a cada uma delas o elemento correspondente de `trimmedParts`. Depois atualize a declaração `return` para retornar um objeto com essas quatro propriedades.

# --hints--

`parseCard` deve retornar um objeto com as propriedades `title`, `author`, `year` e `location`.

```js
const result = parseCard("The Stand | King, Stephen | 1978 | Shelf K2");
assert.isObject(result);
assert.property(result, "title");
assert.property(result, "author");
assert.property(result, "year");
assert.property(result, "location");
```

A propriedade `title` do objeto retornado por `parseCard("The Stand | King, Stephen | 1978 | Shelf K2")` deve ter o valor `"The Stand"`.

```js
assert.strictEqual(parseCard("The Stand | King, Stephen | 1978 | Shelf K2").title, "The Stand");
```

A propriedade `author` do objeto retornado por `parseCard("The Stand | King, Stephen | 1978 | Shelf K2")` deve ter o valor `"King, Stephen"`.

```js
assert.strictEqual(parseCard("The Stand | King, Stephen | 1978 | Shelf K2").author, "King, Stephen");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
--fcc-editable-region--

  return trimmedParts;
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
