---
id: 699d20623e57e8ffee885451
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

A função `parseCard` vai receber uma string de catálogo bruta e retornar um objeto estruturado. Antes de construir sua função, configure uma forma de ver a saída dela enquanto trabalha.

Faça a chamada `parseCard` com `rawCatalogCards[2]` como argumento, atribua o resultado a uma variável chamada `cardResult` e registre `cardResult` no console.

No momento `parseCard` não tem corpo, então retorna `undefined`. Conforme você for adicionando código à função nos próximos passos, a saída registrada mostrará exatamente o que a função está retornando em cada etapa.

# --hints--

Você deve chamar `parseCard` com `rawCatalogCards[2]` como argumento.

```js
assert.match(__helpers.removeJSComments(code), /parseCard\s*\(\s*rawCatalogCards\s*\[\s*2\s*\]\s*\)/);
```

Você deve atribuir o resultado da chamada `parseCard` com `rawCatalogCards[2]` a uma variável chamada `cardResult`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+cardResult\s*=\s*parseCard\s*\(\s*rawCatalogCards\s*\[\s*2\s*\]\s*\)\s*;?/);
```

Você deve registrar `cardResult` no console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*cardResult\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {}

--fcc-editable-region--

--fcc-editable-region--
```
