---
id: 699d20623e57e8ffee885458
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

O campo year chega como uma string, mas você precisa dele como número para um cálculo posterior. Use `parseInt()` para convertê-lo.

O year também pode estar em branco, então você precisa de um fallback. O operador ternário `condition ? valueIfTrue : valueIfFalse` retorna um dos dois valores com base em uma condição.

Atualize a propriedade `year` na instrução return para:

```js
year: year ? parseInt(year) : "Unknown"
```

# --hints--

Quando chamado com um card que tem um year válido, a propriedade `year` do objeto retornado deve ser um número, não uma string.

```js
assert.isNumber(parseCard("The Shining | King, Stephen | 1977 | Shelf K1").year);
```

Quando chamado com `"The Shining | King, Stephen | 1977 | Shelf K1"`, a propriedade `year` deve ser `1977`.

```js
assert.strictEqual(parseCard("The Shining | King, Stephen | 1977 | Shelf K1").year, 1977);
```

Quando chamado com uma string que tem o campo year em branco, a propriedade `year` deve ser `"Unknown"`.

```js
assert.strictEqual(parseCard("X | A |  | Y").year, "Unknown");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
--fcc-editable-region--
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year,
    location: location || "Unknown"
  };
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
