---
id: 699d1f8e51fcdc8230f779fb
title: passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Nesta oficina, você vai construir um sistema de catálogo de biblioteca patrimonial. Você vai analisar dados brutos de livros, buscar e agrupar entradas, validar a qualidade dos dados e exportar os resultados para diferentes formatos.

Cada ficha do catálogo é armazenada como uma única string no formato `"Title | Author | Year | Location"`, com os campos separados pelo caractere pipe (`|`).

Você vai usar um array dessas strings durante toda a oficina para construir e testar suas funções de catálogo. Crie um array chamado `rawCatalogCards` e adicione as duas strings a seguir nele:

```js
"From a Buick 8 | King, Stephen | 2002 | Shelf K7",
"The Shining | King, Stephen | 1977 | Shelf K1"
```

# --hints--

Você deve declarar uma variável chamada `rawCatalogCards`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+rawCatalogCards\s*=/);
```

`rawCatalogCards` deve ser um array.

```js
assert.isArray(rawCatalogCards);
```

O primeiro elemento de `rawCatalogCards` deve ser `"From a Buick 8 | King, Stephen | 2002 | Shelf K7"`.

```js
assert.strictEqual(rawCatalogCards[0], "From a Buick 8 | King, Stephen | 2002 | Shelf K7");
```

O segundo elemento de `rawCatalogCards` deve ser `"The Shining | King, Stephen | 1977 | Shelf K1"`.

```js
assert.strictEqual(rawCatalogCards[1], "The Shining | King, Stephen | 1977 | Shelf K1");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
