---
id: 69d5e6f7a8b9c0d1e2f3a4b5
title: passo 22
challengeType: 1
dashedName: step-22
---

# --description--

Agora preencha `renderEntry`. Declare uma constante fallback para cada campo: `const title = entry.title || "Unknown"`, e aplique o mesmo padrão para `author`, `year` e `location`.

Retorne um template literal multilinha que formate a entrada como um cartão de biblioteca. Use `"-".repeat(25)` para a borda superior e inferior, e inclua linhas rotuladas para `Title`, `Author`, `Year` e `Location`. Por exemplo, uma linha rotulada para o título ficaria assim: `Title: ${title}`.

# --hints--

`renderEntry` deve retornar uma string.

```js
assert.isString(renderEntry(catalog[0]));
```

Quando chamado com `catalog[0]`, a string retornada deve conter `"From a Buick 8"`.

```js
assert.include(renderEntry(catalog[0]), "From a Buick 8");
```

Quando chamado com `catalog[0]`, a string retornada deve conter `"King, Stephen"`.

```js
assert.include(renderEntry(catalog[0]), "King, Stephen");
```

Quando chamado com `catalog[0]`, a string retornada deve conter o ano `2002`.

```js
assert.include(renderEntry(catalog[0]), "2002");
```

Quando chamado com `catalog[0]`, a string retornada deve conter `"Shelf K7"`.

```js
assert.include(renderEntry(catalog[0]), "Shelf K7");
```

A string retornada deve incluir uma borda superior e inferior usando `"-".repeat(25)`.

```js
assert.isAbove(renderEntry(catalog[0]).split("-".repeat(25)).length, 2);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

function parseCatalog(rawCards) {
  const catalog = [];
  for (let i = 0; i < rawCards.length; i++) {
    catalog.push(parseCard(rawCards[i]));
  }
  return catalog;
}

const catalog = parseCatalog(rawCatalogCards);

function findByAuthor(catalog, author) {
  const searchTerm = author.toLowerCase();
  const results = [];
  for (let i = 0; i < catalog.length; i++) {
    if (catalog[i].author.toLowerCase().includes(searchTerm)) {
      results.push(catalog[i]);
    }
  }
  return results;
}

function groupByDecade(catalog) {
  const grouped = {};
  for (let i = 0; i < catalog.length; i++) {
    const book = catalog[i];
    if (book.year === "Unknown") {
      if (!grouped["Unknown"]) {
        grouped["Unknown"] = [];
      }
      grouped["Unknown"].push(book);
      continue;
    }
    const decade = Math.floor(book.year / 10) * 10;
    const decadeKey = `${decade}s`;
    if (!grouped[decadeKey]) {
      grouped[decadeKey] = [];
    }
    grouped[decadeKey].push(book);
  }
  return grouped;
}

const byDecade = groupByDecade(catalog);

function renderEntry(entry) {
--fcc-editable-region--

--fcc-editable-region--
}

console.log(renderEntry(catalog[0]));
```
