---
id: 669e81368e52b3a5c35a2dc5
title: Jenga ukurasa wa mkusanyiko wa video
challengeType: 25
dashedName: build-a-video-compilation-page
demoType: onClick
---

# --description--

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kipengele cha `main` kama mtoto pekee wa kipengele cha `body`.
1. Unapaswa kuwa na kipengele cha `h1` chenye mada ya ukurasa wako.
1. Unapaswa kuwa na aya inayotambulisha mada ya ukurasa wako chini ya kipengele chako cha `h1`.
1. Unapaswa kuwa na vipengele vitatu vya `section` chini ya aya yako ya kwanza.
1. Kila sehemu inapaswa kuwa na kipengele cha `h2`, aya, na kipengele cha `iframe`, kwa mpangilio huu.
1. Vipengele vitatu vya `iframe` vinapaswa kuwa na sifa ya `src` iliyowekwa kwa video halali.
1. Kila kipengele cha `iframe` pia kinapaswa kuwa na sifa ya `title` kuelezea maudhui yaliyowekwa, na sifa za `height` na `width` kuweka kipengele kwa ukubwa unaofaa.

# --hints--

Unapaswa kuwa na kipengele cha `main` ndani ya kipengele chako cha `body`.

```js
assert.equal(document.querySelector('body :first-child').tagName, 'MAIN');
```

Kipengele chako cha `main` kinapaswa kuwa mtoto pekee wa kipengele cha `body`.

```js
assert.equal(document.querySelector('body').children.length, 1);
```

Unapaswa kuwa na kipengele cha `h1` chenye mada ya ukurasa wako ndani ya kipengele cha `main`.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Unapaswa kuwa na aya inayotambulisha mada ya ukurasa wako chini ya kipengele chako cha `h1`.

```js
const p = document.querySelectorAll('p')[0];
assert.equal(p?.previousElementSibling?.tagName, 'H1');
assert.isAbove(p?.innerText.length, 0);
```

Unapaswa kuwa na vipengele vitatu vya `section` chini ya kipengele chako cha `p` cha kwanza.

```js
assert.equal(document.querySelectorAll('section')?.length, 3);
assert.equal(document.querySelectorAll('section')[0]?.previousElementSibling?.tagName, 'P')
```

Kila kipengele cha `section` kinapaswa kuanza na kipengele cha `h2` kinachotumika kama kichwa cha sehemu hiyo.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[0]?.tagName, 'H2');
    assert.isAbove(section.children[0]?.innerText.length, 0);
}
```

Kila kipengele cha `section` kinapaswa kuwa na kipengele cha `p` kutambulisha maudhui ya video kama mtoto wake wa pili.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[1]?.tagName, 'P');
    assert.isAbove(section.children[1]?.innerText.length, 0);
}
```

Kila kipengele cha `section` kinapaswa kuwa na kipengele cha `iframe` kama mtoto wake wa tatu.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[2]?.tagName, 'IFRAME');
}
```

Kipengele chako cha `iframe` cha kwanza kinapaswa kuwa na sifa ya `src` iliyowekwa kwa video halali.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.src.length, 0)
```

Kipengele chako cha `iframe` cha kwanza kinapaswa kuwa na sifa ya `title` kuelezea maudhui yaliyowekwa.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.title.length, 0)
```

Kipengele chako cha `iframe` cha kwanza kinapaswa kuwa na sifa ya `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.height.length, 0)
```

Kipengele chako cha `iframe` cha kwanza kinapaswa kuwa na sifa ya `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.width.length, 0)
```

Kipengele chako cha `iframe` cha pili kinapaswa kuwa na sifa ya `src` iliyowekwa kwa video halali.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.src.length, 0)
```

Kipengele chako cha `iframe` cha pili kinapaswa kuwa na sifa ya `title` kuelezea maudhui yaliyowekwa.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.title.length, 0)
```

Kipengele chako cha `iframe` cha pili kinapaswa kuwa na sifa ya `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.height.length, 0)
```

Kipengele chako cha `iframe` cha pili kinapaswa kuwa na sifa ya `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.width.length, 0)
```

Kipengele chako cha `iframe` cha tatu kinapaswa kuwa na sifa ya `src` iliyowekwa kwa video halali.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.src.length, 0)
```

Kipengele chako cha `iframe` cha tatu kinapaswa kuwa na sifa ya `title` kuelezea maudhui yaliyowekwa.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.title.length, 0)
```

Kipengele chako cha `iframe` cha tatu kinapaswa kuwa na sifa ya `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.height.length, 0)
```

Kipengele chako cha `iframe` cha tatu kinapaswa kuwa na sifa ya `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.width.length, 0)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>
    <main>
        <h1>Front-End Web Development</h1>
        <p>Front-End Web Development involves designing and building user interfaces of websites using HTML, CSS, and
            JavaScript. It focuses on creating visually appealing, interactive, and user-friendly websites. Front-end
            developers ensure accessibility, performance optimization, and
            adherence to web standards, blending creativity with technical expertise.</p>
        <section>
            <h2>HTML</h2>
            <p>HTML, or HyperText Markup Language, forms the backbone of web pages by defining their basic structure and
                layout. It uses a series of elements to display text, images, lists, tables, and other static content on
                a web page.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/GDGejH3SDNQ?si=KJYLgcz4kyyroYMB"
                title="html and coding introduction video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>CSS</h2>
            <p>CSS, or Cascading Style Sheets, is used to control the visual presentation of HTML elements on a web
                page. It allows for customized styling, including colors, fonts, layouts, and spacing, making websites
                visually appealing. CSS also supports responsive design, ensuring web pages look good on
                all devices.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/OXGznpKZ_sA?si=s3KDSZvrhU_PU9XL"
                title="css tutorial video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>JavaScript</h2>
            <p>JavaScript is a programming language that adds interactivity to static web pages, enabling dynamic
                content updates, form validations, animations, and responsive behaviors based on user interactions and
                events.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/jS4aFq5-91M?si=zKQhHEYwU4tnMmVm"
                title="javascript programming video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
    </main>
</body>

</html>
```
