---
id: 66ec4c8e9878d8441956516f
title: Erstellen Sie eine Buchkatalog-Tabelle
challengeType: 25
dashedName: build-a-book-catalog-table
demoType: onClick
---

# --description--

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bestehen Sie alle Tests, um die Aufgabe abzuschließen.

**User Stories:**

1. Sie sollten ein `table`-Element erstellen, das Buchinformationen auflistet.
1. Ihre Tabelle sollte ein Tabellenkopf-Element mit einer Zeile enthalten.
1. Die Zeile im Tabellenkopf-Element sollte vier Tabellenkopf-Elemente enthalten, mit dem Text `Title`, `Author`, `Genre` und `Publication Year` in dieser Reihenfolge.
1. Ihre Tabelle sollte ein Tabellenkörper-Element mit mindestens fünf Zeilen enthalten.
1. Jede Zeile im Tabellenkörper sollte vier Tabellendaten-Elemente enthalten, die den Titel, Autor, das Genre und das Veröffentlichungsjahr des Buchs anzeigen.
1. Ihre Tabelle sollte ein Tabellenfuß-Element mit einer Zeile enthalten.
1. Die Zeile im Tabellenfuß-Element sollte ein Tabellendaten-Element enthalten, das sich über vier Spalten erstreckt und den Text `Total Books: N` enthält, wobei `N` durch die Anzahl der Bücher in Ihrer Tabelle ersetzt werden sollte.

# --hints--

Sie sollten ein `table`-Element haben.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

Sie sollten ein `thead`-Element innerhalb Ihres `table`-Elements haben.

```js
assert.lengthOf(document.querySelectorAll('table thead'), 1);
```

Sie sollten ein `tr`-Element innerhalb Ihres `thead`-Elements haben.

```js
assert.lengthOf(document.querySelectorAll('thead tr'), 1);
```

Sie sollten vier `th`-Elemente innerhalb der Zeile Ihres `thead`-Elements haben.

```js
assert.lengthOf(document.querySelectorAll('thead tr th'), 4);
```

Ihre vier `th`-Elemente sollten den Text `Title`, `Author`, `Genre` und `Publication Year` in dieser Reihenfolge enthalten.

```js
const ths = document.querySelectorAll('thead tr th');
assert.equal(ths[0]?.textContent.trim(), 'Title');
assert.equal(ths[1]?.textContent.trim(), 'Author');
assert.equal(ths[2]?.textContent.trim(), 'Genre');
assert.equal(ths[3]?.textContent.trim(), 'Publication Year');
```

Sie sollten ein `tbody`-Element innerhalb Ihres `table`-Elements haben.

```js
assert.lengthOf(document.querySelectorAll('table tbody'), 1);
```

Ihr `tbody`-Element sollte mindestens fünf Zeilen enthalten.

```js
assert.isAtLeast(document.querySelectorAll('tbody tr').length, 5);
```

Jede Zeile in Ihrem `tbody`-Element sollte genau vier `td`-Elemente als Kinder haben.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isAtLeast(rows.length, 1);

rows.forEach(row => {
  assert.lengthOf(row.querySelectorAll('td'), 4);
  assert.lengthOf(row.children, 4);
});
```

Jedes `td`-Element im Tabellenkörper sollte Text mit Buchinformationen enthalten.

```js
const tds = document.querySelectorAll('tbody tr td');
assert.isAtLeast(tds.length, 1);

tds.forEach(td => {
  assert.isAtLeast(td.textContent.trim().length, 1);
});
```

Sie sollten ein `tfoot`-Element innerhalb Ihres `table`-Elements haben.

```js
assert.lengthOf(document.querySelectorAll('table tfoot'), 1);
```

Sie sollten genau ein `tr`-Element in Ihrem `tfoot`-Element haben.

```js
assert.lengthOf(document.querySelectorAll('tfoot tr'), 1);
```

Das `td`-Element in der Zeile Ihres `tfoot`-Elements sollte sein `colspan`-Attribut auf `4` gesetzt haben.

```js
assert.equal(document.querySelector('tfoot tr td')?.colSpan, 4);
```

Das `td`-Element in der Zeile Ihres `tfoot`-Elements sollte den Text `Total Books: N` enthalten, wobei `N` die Anzahl der Bücher in Ihrer Tabelle ist.

```js
const numberOfBooks = document.querySelectorAll('tbody tr')?.length;
assert.equal(document.querySelector('tfoot tr td').textContent.trim(), `Total Books: ${numberOfBooks}`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Book Catalog</title>
</head>

<body>
</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Book Catalog</title>
</head>
<body>
    <h1>Book Catalog</h1>
  <table>
    <thead>
      <tr>
        <th>Title</th>
        <th>Author</th>
        <th>Genre</th>
        <th>Publication Year</th>
      </tr>
    </thead>
    <tbody>
      <tr>
        <td>How to Contribute to Open-Source Projects – A Handbook for Beginners</td>
        <td>Hillary Nyakundi</td>
        <td>Open Source</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>Learn Linux for Beginners: From Basics to Advanced Techniques</td>
        <td>Zaira Hira</td>
        <td>Linux</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>How to Learn to Code and Get a Developer Job</td>
        <td>Quincy Larson</td>
        <td>Learn To Code</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>The Regular Expressions Book – RegEx for JavaScript Developers</td>
        <td>Kolade Chris</td>
        <td>Regular Expressions</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>The Python Code Example Handbook</td>
        <td>Farhan Hasin Chowdhury</td>
        <td>Python</td>
        <td>2023</td>
      </tr>
    </tbody>
    <tfoot>
      <tr>
        <td colspan="4">Total Books: 5</td>
      </tr>
    </tfoot>
  </table>
</body>
</html>
```
