---
id: 66db529d37ad966480ebb633
title: ランチピッカープログラムを作ろう
challengeType: 26
dashedName: build-a-lunch-picker-program
---

# --description--

このラボでは、ランチの選択肢を管理するプログラムを作成します。ランチの配列を扱い、配列にアイテムを追加・削除し、ランチの選択肢をランダムに選ぶ処理を実装します。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `lunches` という変数を宣言し、ランチのアイテムをストアするための空の配列を代入してください。

2. 配列を最初の引数、文字列を2番目の引数に取る関数 `addLunchToEnd` を作成してください。この関数は以下を行います:
    - 文字列を配列の末尾に追加する。
    - `[Lunch Item] added to the end of the lunch menu.` に渡された文字列を使って、文字列 `[Lunch Item]` をコンソールにログ出力する。
    - 更新された配列を返す。

3. 配列を最初の引数、文字列を2番目の引数に取る関数 `addLunchToStart` を作成してください。この関数は以下を行います:
    - 文字列を配列の先頭に追加する。
    - `[Lunch Item] added to the start of the lunch menu.` に渡された文字列を使って、文字列 `[Lunch Item]` をコンソールにログ出力する。
    - 更新された配列を返す。

4. 配列を引数に取る関数 `removeLastLunch` を作成してください。この関数は以下を行います:
   - 配列の最後の要素を削除する。
   - 削除が成功した場合、削除された要素 `[Lunch Item] removed from the end of the lunch menu.` を使って文字列 `[Lunch Item]` をコンソールにログ出力する。
   - 配列が空の場合は、文字列 `"No lunches to remove."` をコンソールにログ出力する。
   - 更新された配列を返す。

5. 配列を引数に取る関数 `removeFirstLunch` を作成してください。この関数は以下を行います:
   - 配列の最初の要素を削除する。
   - 削除が成功した場合、削除された要素 `[Lunch Item] removed from the start of the lunch menu.` を使って文字列 `[Lunch Item]` をコンソールにログ出力する。
   - 配列が空の場合は、文字列 `"No lunches to remove."` をコンソールにログ出力する。
   - 更新された配列を返す。

6. 配列を引数に取る関数 `getRandomLunch` を作成してください。この関数は以下を行います:
   - 配列からランダムな要素を選ぶ。
   - 成功した場合、配列のランダムな要素 `Randomly selected lunch: [Lunch Item]` を使って文字列 `[Lunch Item]` をコンソールにログ出力する。
   - 配列が空の場合は、文字列 `"No lunches available."` をコンソールにログ出力する。

7. 配列を引数に取る関数 `showLunchMenu` を作成してください。この関数は以下を行います:
   - 配列に要素がある場合、配列の各要素 `Menu items: [Lunch Item], [Lunch Item]...` を順番に使って文字列 `[Lunch item]` をコンソールにログ出力する。
   - 配列が空の場合は、文字列 `"The menu is empty."` をコンソールにログ出力する。

# --hints--

ランチのアイテムをストアするために、`lunches` という変数を宣言し、空の配列を代入してください。

```js
assert.isArray(lunches);
```

`addLunchToEnd` という関数を定義してください。

```js
assert.isFunction(addLunchToEnd);
```

関数 `addLunchToEnd` は2つのパラメータを持つべきです。

```js
assert.lengthOf(addLunchToEnd, 2);
```

`addLunchToEnd(lunches, "Tacos")` は文字列 `"Tacos added to the end of the lunch menu."` をコンソールにログ出力するべきです。

```js
const tempArr = [];
const temp = console.log;
const testLunches = []
try {
  console.log = obj => tempArr.push(obj);
  addLunchToEnd(testLunches, "Tacos");
  assert.strictEqual(tempArr[0], "Tacos added to the end of the lunch menu.");
  addLunchToEnd(testLunches, "Pizza");
  assert.strictEqual(tempArr[1], "Pizza added to the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToEnd(["Pizza", "Tacos"], "Burger")` は `["Pizza", "Tacos", "Burger"]` を返すべきです。

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToEnd(["Pizza", "Tacos"], "Burger"), ["Pizza", "Tacos", "Burger"]);
  assert.deepStrictEqual(addLunchToEnd(["Fries"], "Tacos"), ["Fries", "Tacos"]);
} finally {
  console.log = temp;
}
```

`addLunchToStart` という関数を定義してください。

```js
assert.isFunction(addLunchToStart);
```

関数 `addLunchToStart` は2つのパラメータを持つべきです。

```js
assert.lengthOf(addLunchToStart, 2);
```

`addLunchToStart(lunches, "Sushi")` は文字列 `"Sushi added to the start of the lunch menu."` をコンソールにログ出力するべきです。

```js
const tempArr = [];
const testLunches = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  addLunchToStart(testLunches, "Sushi");
  assert.strictEqual(tempArr[0], "Sushi added to the start of the lunch menu.");  
  addLunchToStart(testLunches, "Burger");
  assert.strictEqual(tempArr[1], "Burger added to the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToStart(["Burger", "Sushi"], "Pizza")` は `["Pizza", "Burger", "Sushi"]` を返すべきです。

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToStart(["Burger", "Sushi"], "Pizza"), ["Pizza", "Burger", "Sushi"]);
  assert.deepStrictEqual(addLunchToStart(["Noodles"], "Sushi"), ["Sushi", "Noodles"]);
} finally {
  console.log = temp;
}
```

`removeLastLunch` という関数を定義してください。

```js
assert.isFunction(removeLastLunch)
```

関数 `removeLastLunch` は1つのパラメータを持つべきです。

```js
assert.lengthOf(removeLastLunch, 1);
```

入力配列が空の場合、関数 `removeLastLunch` は文字列 `"No lunches to remove."` をコンソールにログ出力するべきです。

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Stew", "Soup", "Toast"])` は文字列 `"Toast removed from the end of the lunch menu."` をコンソールにログ出力するべきです。

```js
let testLunches = ["Stew", "Soup", "Toast"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "Toast removed from the end of the lunch menu.");
  testLunches = ["Rice", "Pizza"];
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[1], "Pizza removed from the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Sushi", "Pizza", "Noodles"])` は `["Sushi", "Pizza"]` を返すべきです。

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeLastLunch(["Sushi", "Pizza", "Noodles"]), ["Sushi", "Pizza"]);
  assert.deepStrictEqual(removeLastLunch(["Pizza", "Burger"]), ["Pizza"]);
} finally {
  console.log = temp;
}
```

`removeFirstLunch` という関数を定義してください。

```js
assert.isFunction(removeFirstLunch);
```

関数 `removeFirstLunch` は1つのパラメータを持つべきです。

```js
assert.lengthOf(removeFirstLunch, 1);
```

入力配列が空の場合、関数 `removeFirstLunch` は文字列 `"No lunches to remove."` をコンソールにログ出力するべきです。

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Salad", "Eggs", "Cheese"])` は文字列 `"Salad removed from the start of the lunch menu."` をコンソールにログ出力するべきです。

```js
let testLunches = ["Salad", "Eggs", "Cheese"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "Salad removed from the start of the lunch menu.");
  testLunches = ["Stew", "Soup", "Toast"]
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[1], "Stew removed from the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Sushi", "Pizza", "Burger"])` は `["Pizza", "Burger"]` を返すべきです。

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeFirstLunch(["Sushi", "Pizza", "Burger"]), ["Pizza", "Burger"]);
  assert.deepStrictEqual(removeFirstLunch(["Pizza", "Burger"]), ["Burger"]);
} finally {
  console.log = temp;
}
```

`addLunchToEnd`、`addLunchToStart`、`removeLastLunch`、および `removeFirstLunch` は、更新後に引数として渡された同じ配列を返すべきです。

```js
const temp = console.log;
console.log = () => {}
const testLunches = ["Spam"];
try {
  assert.equal(addLunchToEnd(testLunches, "Pizza"), testLunches);
  assert.equal(addLunchToStart(testLunches, "Caviar"), testLunches);
  assert.equal(removeLastLunch(testLunches), testLunches);
  assert.equal(removeFirstLunch(testLunches), testLunches);
} finally {
  console.log = temp;
}
```

`getRandomLunch` という関数を定義してください。

```js
assert.isFunction(getRandomLunch);
```

関数 `getRandomLunch` は1つのパラメータを持つべきです。

```js
assert.lengthOf(getRandomLunch, 1);
```

入力配列が空の場合、関数 `getRandomLunch` は文字列 `"No lunches available."` をコンソールにログ出力するべきです。

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  getRandomLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches available.");
} finally {
  console.log = temp;
}
```

入力配列が空でない場合、関数 `getRandomLunch` は `Randomly selected lunch: [Lunch Item]` の形式の文字列をコンソールにログ出力するべきです。

```js
const testLunches = ["Fish", "Fries", "Roast"];
const tempRandom = Math.random;
const tempArr = [];
const temp = console.log;

try {
  console.log = obj => tempArr.push(obj);
  // check that it correctly outputs the first item
  Math.random = () => 0;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[0], `Randomly selected lunch: ${testLunches[0]}`);

  // Second item
  Math.random = () => 0.4;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[1], `Randomly selected lunch: ${testLunches[1]}`);

  // Third item
  Math.random = () => 0.8;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[2], `Randomly selected lunch: ${testLunches[2]}`);

} finally {
  Math.random = tempRandom;
  console.log = temp;
}
```

関数 `getRandomLunch` は引数として渡された配列を変更してはいけません。

```js
const testLunches = ["Fish", "Fries", "Roast"];
let originalLunches = [...testLunches];
const tempRandom = Math.random;

try {
  Math.random = () => 0;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.4;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.8;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  const emptyLunches = [];
  let originalEmpty = [...emptyLunches];
  getRandomLunch(emptyLunches);
  assert.sameOrderedMembers(emptyLunches, originalEmpty);

} finally {
  Math.random = tempRandom;
}
```

`showLunchMenu` という関数を定義してください。

```js
assert.isFunction(showLunchMenu);
```

関数 `showLunchMenu` は1つのパラメータを持つべきです。

```js
assert.lengthOf(showLunchMenu, 1);
```

入力配列が空の場合、関数 `showLunchMenu` は文字列 `"The menu is empty."` をコンソールにログ出力するべきです。

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], "The menu is empty.");
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Greens", "Corns", "Beans"])` は文字列 `"Menu items: Greens, Corns, Beans"` をコンソールにログ出力するべきです。

```js
const testLunches = ["Greens", "Corns", "Beans"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Pizza", "Burger", "Fries", "Salad"])` は文字列 `"Menu items: Pizza, Burger, Fries, Salad"` をコンソールにログ出力するべきです。

```js
const testLunches = ["Pizza", "Burger", "Fries", "Salad"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const lunches = [];

// Add a new lunch to the end of the list
function addLunchToEnd(lunches, newLunch) {
  lunches.push(newLunch);
  console.log(`${newLunch} added to the end of the lunch menu.`);
  return lunches;
}

// Add a new lunch to the start of the list
function addLunchToStart(lunches, newLunch) {
  lunches.unshift(newLunch);
  console.log(`${newLunch} added to the start of the lunch menu.`);
  return lunches;
}

// Remove the last lunch from the list
function removeLastLunch(lunches) {
  const removedLunch = lunches.pop();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the end of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Remove the first lunch from the list
function removeFirstLunch(lunches) {
  const removedLunch = lunches.shift();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the start of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Function to get a random lunch from the list
function getRandomLunch(lunches) {
  if (lunches.length === 0) {
    console.log("No lunches available.");
  }
  else {
    const randomIndex = Math.floor(Math.random() * lunches.length);
    const randomLunch = lunches[randomIndex];
    console.log(`Randomly selected lunch: ${randomLunch}`);
  }
}

// Function to display all the lunches in the list
function showLunchMenu(lunches) {
  if (lunches.length === 0) {
    console.log("The menu is empty.");
  }
  else {
    console.log(`Menu items: ${lunches.join(', ')}`);
  }
}

addLunchToEnd(lunches, "Tacos");
addLunchToStart(lunches, "Sushi");
removeLastLunch(lunches);
removeFirstLunch(lunches);
getRandomLunch(lunches);
showLunchMenu(lunches);

```
