---
id: 694bf6eb530e19a9c48c59f1
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

يبدو أن عنوانك غير مكتمل الآن. تريد أيضًا إضافة رقم الشقة التي يعيش فيها الموظف، لذلك يجب تعديل المتغير.

عندما تريد إضافة محتوى إلى نهاية متغير سلسلة نصية موجود، يمكنك استخدام معامل **التعيين المعزز** `+=`. هذا أقصر من كتابة `var = var + 'new text'`. على سبيل المثال:

```py
greeting = 'Hello'
greeting += ' World'
print(greeting) # Hello World
```

تذكر أن السلاسل النصية غير قابلة للتغيير، لذلك هذه العملية لا تغير السلسلة الأصلية. بل تُنشئ سلسلة جديدة وتُعيد تعيينها إلى المتغير.

استخدم معامل `+=` لإضافة السلسلة النصية `, Apartment 4B` إلى متغير `address` الخاص بك.

# --hints--

يجب استخدام معامل `+=` لإضافة السلسلة النصية `, Apartment 4B` إلى المتغير `address`.

```js
({ test: () => assert(runPython(`_Node(_code).find_aug_variable("address").is_equivalent("address += ', Apartment 4B'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
address = '123 Main Street'  
--fcc-editable-region--  

--fcc-editable-region-- 
print(address)  
```
