---
id: 694bf6eb530e19a9c48c59f7
title: الخطوة 18
challengeType: 20
dashedName: step-18
---

# --description--

يمكنك أيضًا استخدام الأعداد السالبة لاقتطاع جزء من نهاية سلسلة نصية. على سبيل المثال، تشير `-1` إلى الحرف الأخير، و`-2` إلى الحرف قبل الأخير، وهكذا.

للحصول على آخر ثلاثة أحرف من سلسلة نصية، يمكنك استخدام `string[-3:]`. لاحظ كيف تم حذف معلمة التوقف بعد النقطتين. هذا يعني أن الاقتطاع يجب أن يستمر حتى نهاية السلسلة النصية.

أنشئ متغيرًا باسم `last_three`. استخدم الفهرسة السالبة لاستخراج آخر ثلاثة أحرف من `employee_code` (التي تمثل رقم الهوية). أخيرًا، اطبع `last_three` على الطرفية.

بهذا، تكتمل ورشة عمل بطاقة الموظف.

# --hints--

يجب أن يكون لديك متغير باسم `last_three`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("last_three")`)) })
```

يجب أن تقتطع آخر ثلاثة أحرف من `employee_code` باستخدام الفهرسة السالبة وتعيينها إلى متغير `last_three` الخاص بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_three").is_equivalent("last_three = employee_code[-3:]")`)) })
```

يجب أن تطبع `last_three` على الطرفية.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_three)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
print(year_code)
initials = employee_code[9:11]
print(initials)
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
initials = employee_code[9:11]
print(year_code)
print(initials)
last_three = employee_code[-3:]
print(last_three)
```
