---
id: 694bf6eb530e19a9c48c59f2
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

أكمل الجملة الآن بتحديث تعيين `employee_info` ليشمل أيضًا ربط السلسلة النصية ` years old` في النهاية، باستخدام العامل `+`. تذكر تضمين مسافة في بداية السلسلة النصية.

أخيرًا، اطبع `employee_info`.

# --hints--

يجب ربط `' years old'` في نهاية المتغير `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age) + ' years old'")`)) })
```

يجب طباعة `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_info)")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + str(employee_age)

--fcc-editable-region-- 
```
