---
id: 694bf6eb530e19a9c48c59f3
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

الآن ستستخدم الدالة `str()` مرة أخرى. تمامًا كما في حالة age، يجب تحويل أي متغير رقمي إلى سلسلة نصية قبل دمجه مع نص آخر.

أنشئ متغيرًا باسم `experience_years` وعيّن له العدد الصحيح `5`.

ثم، أنشئ متغيرًا باسم `experience_info`. عيّن له سلسلة نصية تتكون من دمج `'Experience: '`، ومتغير `experience_years` (بعد تحويله إلى سلسلة نصية)، و`' years'`، باستخدام عامل `+`. اطبع النتيجة في الطرفية.

# --hints--

يجب أن يكون لديك متغير باسم `experience_years`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_years")`)) })
```

يجب تعيين العدد الصحيح `5` إلى متغير `experience_years` الخاص بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_years").is_equivalent("experience_years = 5")`)) })
```

يجب أن يكون لديك متغير باسم `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_info")`)) })
```

يجب تعيين سلسلة نصية تتكون من دمج `'Experience: '`, `str(experience_years)`, و`' years'` إلى متغير `experience_info` الخاص بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_info").is_equivalent("experience_info = 'Experience: ' + str(experience_years) + ' years'")`)) })
```

يجب طباعة `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(experience_info)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
--fcc-editable-region--

--fcc-editable-region--
```
