---
id: 694bf3af05d10fd89dab05cc
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

في Python، يمكنك دمج السلاسل النصية باستخدام العامل `+`. هذا يُسمى الربط. إليك مثالًا:

```python
greeting = 'Hello' + 'World'
print(greeting) # Output: HelloWorld
```

أنشئ متغيرًا باسم `full_name` بربط `first_name` و `last_name`. ثم اطبع `full_name`.

# --hints--

يجب أن يكون لديك متغير باسم `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("full_name")`)) })
```

يجب أن تُعيّن سلسلة نصية ناتجة عن ربط `first_name` و `last_name` إلى متغير `full_name` الخاص بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + last_name")`)) })
```

يجب أن تطبع متغير `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(full_name)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--

--fcc-editable-region--
```
