---
id: 694bf6eb530e19a9c48c59f4
title: الخطوة 13
challengeType: 20
dashedName: step-13
---

# --description--

ربط العديد من السلاسل النصية باستخدام `+` وتحويل الأرقام باستخدام `str()` قد يصبح فوضويًا وصعب القراءة.

قدم Python 3.6 **سلاسل f** لحل هذه المشكلة. بإضافة الحرف `f` قبل علامة الاقتباس الافتتاحية، يمكنك وضع المتغيرات والتعابير داخل حقول الاستبدال الممثلة بأقواس متعرجة `{}`. على سبيل المثال:

```python
name = 'John'
print(f'Hello {name}') # Output: Hello John
```

أنشئ متغيرًا باسم `employee_card` وعيّن له سلسلة f تعرض `Employee:` متبوعة بمسافة وقيمة المتغير `full_name`.

# --hints--

يجب أن يكون لديك متغير باسم `employee_card`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_card")`)) })
```

يجب أن يحتوي متغير `employee_card` على القيمة `f'Employee: {full_name}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name}'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

--fcc-editable-region--
```
