---
id: 694bf6eb530e19a9c48c59f5
title: الخطوة 16
challengeType: 20
dashedName: step-16
---

# --description--

عند العمل مع السلاسل النصية، ستحتاج غالبًا إلى استخراج جزء محدد من السلسلة. يُسمى هذا **التقطيع**.

الصيغة النحوية هي `string[start:stop]`، حيث:

* `start` هو الفهرس الذي يبدأ منه التقطيع (**شامل**).
* `stop` هو الفهرس الذي ينتهي عنده التقطيع (**حصري**).

على سبيل المثال، إذا كانت `text = 'Python'`، فإن `text[0:2]` تعطي `'Py'`.

عرّف `employee_code` كسلسلة `'DEV-2026-JD-001'`. بعد ذلك، أنشئ متغيرًا باسم `department` وقم بتعيين التقطيع من `employee_code` من الفهرس `0` إلى `3`. ثم اطبع `department` على الطرفية.

# --hints--
يجب أن يكون لديك متغير باسم `employee_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_code")`)) })
```

يجب تعيين السلسلة `DEV-2026-JD-001` إلى متغير `employee_code` الخاص بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_code").is_equivalent("employee_code = 'DEV-2026-JD-001'")`)) })
```

يجب أن يكون لديك متغير باسم `department`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("department")`)) })
```

يجب تقطيع أول ثلاثة أحرف من `employee_code` وتعيينها إلى متغير `department` الخاص بك.

```js
({ test: () => runPython(`
dep = _Node(_code).find_variable("department")
assert dep.is_equivalent("department = employee_code[0:3]") or dep.is_equivalent("department = employee_code[:3]")
`) })
```

يجب طباعة `department` على الطرفية.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(department)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)

--fcc-editable-region--

--fcc-editable-region--
```
