---
id: 694bf6eb530e19a9c48c59f6
title: الخطوة 17
challengeType: 20
dashedName: step-17
---

# --description--

يمكنك اقتطاع جزء من أي مكان في السلسلة النصية، وليس فقط من البداية. وهذا مفيد في حالات كثيرة.

أنشئ متغيرًا باسم `year_code` وقم بتعيين الاقتطاع من `employee_code` من الفهرس `4` إلى `8`. هذا سيستخرج `2026`.

ثم أنشئ متغيرًا باسم `initials` وقم بتعيين الاقتطاع من `employee_code` من الفهرس `9` إلى `11`. هذا سيستخرج `JD`.

أخيرًا، اطبع المتغيرين في الطرفية.

# --hints--

يجب أن يكون لديك متغير باسم `year_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("year_code")`)) })
```

يجب أن تقتطع من `employee_code` من الفهرس `4` إلى `8` وتعيّن النتيجة إلى متغير `year_code`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("year_code").is_equivalent("year_code = employee_code[4:8]")`)) })
```

يجب أن يكون لديك متغير باسم `initials`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("initials")`)) })
```

يجب أن تقتطع من `employee_code` من الفهرس `9` إلى `11` وتعيّن النتيجة إلى متغير `initials`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("initials").is_equivalent("initials = employee_code[9:11]")`)) })
```

يجب أن تطبع `year_code` في الطرفية.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(year_code)")`)) })
```

يجب أن تطبع `initials` في الطرفية.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(initials)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department) 
--fcc-editable-region--  

--fcc-editable-region--  
```
