---
id: 694bf6ea530e19a9c48c59f0
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

إذا جمعت سلسلتين نصيتين مثل `'John' + 'Doe'`، تكون النتيجة `'JohnDoe'` بدون فراغ. لإصلاح ذلك، تحتاج إلى جمع سلسلة نصية تحتوي على فراغ (`' '`) بينهما.

حدّث المتغير `full_name` ليُدمج المتغير `first_name`، ومسافة (`' '`)، ثم المتغير `last_name` بهذا الترتيب، باستخدام عامل الجمع `+`.

# --hints--

يجب إضافة فراغ (`' '`) بين `first_name` و`last_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + ' ' + last_name")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--
full_name = first_name + last_name
--fcc-editable-region--
print(full_name)
```
