---
id: 695c96f82af4fb78e59fa9ed
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

السلاسل النصية هي متتاليات من الأحرف تُستخدم لتخزين بيانات نصية. كما قد تتذكر من الدروس السابقة، يمكنك إنشاء سلسلة نصية عن طريق إحاطة النص بعلامات اقتباس مفردة (`'`) أو علامات اقتباس مزدوجة (`"`). على سبيل المثال:

```python
greeting = 'Hello'
print(greeting) # Output: Hello
```

أنشئ متغيرًا باسم `first_name` يخزن السلسلة النصية `John` ومتغيرًا باسم `last_name` يخزن السلسلة النصية `Doe`. ثم اطبع `first_name` و`last_name`.

# --hints--

يجب أن يكون لديك متغير باسم `first_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("first_name")`)) })
```

يجب تعيين السلسلة النصية `John` إلى متغير `first_name`. تذكر إحاطة النص إما بعلامات اقتباس مفردة أو علامات اقتباس مزدوجة.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("first_name").is_equivalent("first_name = 'John'")`)) })
```

يجب أن يكون لديك متغير باسم `last_name`.

```js  
({ test: () => assert(runPython(`_Node(_code).has_variable("last_name")`)) })
```

يجب تعيين السلسلة النصية `Doe` إلى متغير `last_name`. تذكر إحاطة النص إما بعلامات اقتباس مفردة أو علامات اقتباس مزدوجة.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_name").is_equivalent("last_name = 'Doe'")`)) })
```

يجب طباعة متغير `first_name` في الطرفية.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(first_name)")`)) })
```

يجب طباعة متغير `last_name` في الطرفية.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_name)")`)) })
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
