---
id: 695c9c7c61d446f1451e0ae7
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

بعد ذلك، أنشئ متغيرًا باسم `address` ليخزن عنوان الموظف. عيّن له السلسلة النصية `123 Main Street`، وأخيرًا اطبع `address`.

# --hints--

يجب أن يكون لديك متغير باسم `address`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("address")`)) })
```

يجب أن تعيّن السلسلة النصية `123 Main Street` إلى متغير `address`. تذكر أن تحيط النص بعلامات اقتباس مفردة أو علامات اقتباس مزدوجة.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("address").is_equivalent("address = '123 Main Street'")`)) })
```

يجب أن تطبع المتغير `address`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(address)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
--fcc-editable-region--

--fcc-editable-region--
```

