---
id: 695c9f21a8b1d53cf1b97841
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

الآن أنشئ متغيرًا باسم `employee_age` وقم بتعيين العدد الصحيح `28` له.


# --hints--

يجب أن يكون لديك متغير باسم `employee_age`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_age")`)) })
```

يجب تعيين العدد الصحيح `28` إلى متغير `employee_age`. لا تحيط النص بعلامات اقتباس.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_age").is_equivalent("employee_age = 28")`)) })
```


# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
--fcc-editable-region--

--fcc-editable-region--
```
