---
id: 695df6480de23758af5a29cd
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

الآن، تريد إنشاء سلسلة نصية تعرض عمر الموظف.

ابدأ بإنشاء متغير `employee_info` وقم بتعيين نتيجة دمج باستخدام عامل `+`:

- المتغير `full_name`.
- سلسلة نصية تتكون من الحروف `is` مسبوقة ومتبوعة بمسافة.

# --hints--

يجب أن يكون لديك متغير باسم `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_info")`)) })
```

يجب تعيين `full_name + ' is '` إلى متغير `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is '")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--

--fcc-editable-region--
```
