---
id: 695ea1e55c0e619e597fc5c9
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

كما ترى، أطلق Python الخطأ `TypeError: can only concatenate str (not "int") to str`. يحدث هذا لأن Python لا يسمح بربط النصوص (السلاسل النصية) والأعداد (الأعداد الصحيحة) مباشرةً.

لإصلاح ذلك، يجب تحويل العدد إلى سلسلة نصية أولاً باستخدام الدالة `str()`، التي تُرجع الإصدار النصي من الكائن:

```py
my_num = str(42)
print(type(my_num)) # <class 'str'>
```

حدّث تعيين `employee_info` لتحويل `employee_age` إلى سلسلة نصية باستخدام `str(employee_age)`.

# --hints--

يجب تعديل تعيين المتغير `employee_info` بتحويل `employee_age` إلى سلسلة نصية باستخدام `str(employee_age)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + employee_age
--fcc-editable-region--
```
