---
id: 695fd3dde03273875e5d6709
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

حدّث تعيين `employee_info` ليشمل أيضًا ربط `employee_age` في النهاية، باستخدام العامل `+`.

بعد إتمام ذلك، سترى `TypeError` في الطرفية. في الخطوة التالية، ستعمل على إصلاحه.

# --hints--

يجب أن تحدّث تعيين `employee_info` لربط `employee_age` في النهاية.

```js
const regex = /^employee_info\s*=\s*full_name\s*\+\s*(["']) is \1\s*\+\s*employee_age\s*$/m;  
assert.match(__helpers.python.removeComments(code), regex);  
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is '
--fcc-editable-region--
```
