---
id: 69727ea5c9606bedfe36d3d2
title: الخطوة 15
challengeType: 20
dashedName: step-15
---

# --description--

حان الوقت الآن لإضافة التفاصيل النهائية إلى البطاقة.

عرّف متغيرًا باسم `position` بقيمة السلسلة النصية `Data Analyst` ومتغيرًا باسم `salary` بقيمة العدد الصحيح `75000`.

بعد ذلك، حدّث سلسلة f-string الخاصة بـ `employee_card` لتشمل المنصب والراتب. يجب أن تتبع هذا التنسيق بالضبط: `Employee: [full_name] | Age: [employee_age] | Position: [position] | Salary: $[salary]`. استبدل عناصر النائب بالمتغيرات المقابلة.

أخيرًا، اطبع `employee_card` لرؤية النتيجة.

# --hints--

يجب أن يكون لديك متغير باسم `position`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("position")`)) })
```

يجب تعيين السلسلة النصية `Data Analyst` إلى متغير `position` الخاص بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("position").is_equivalent("position = 'Data Analyst'")`)) })
```

يجب أن يكون لديك متغير باسم `salary`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("salary")`)) })
```

يجب تعيين العدد الصحيح `75000` إلى متغير `salary` الخاص بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

يجب تحديث `employee_card` لاستخدام سلسلة f-string `f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: \${salary}'")`)) })
```

يجب طباعة `employee_card` في الطرفية.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_card)")`))})
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

employee_card = f'Employee: {full_name} | Age: {employee_age}'
--fcc-editable-region--
```
