---
id: 613297a923965e0703b64796
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Vielleicht kennen Sie das `meta`-Element bereits; es wird verwendet, um Informationen über die Seite anzugeben, wie den Titel, die Beschreibung, Schlüsselwörter und den Autor.

Fügen Sie Ihrer Seite ein `meta`-Element mit einem passenden `charset`-Wert hinzu.

Das `charset`-Attribut gibt die Zeichenkodierung der Seite an, und heutzutage wird `UTF-8` von den meisten Browsern als einzige Kodierung unterstützt.

# --hints--

Sie sollten ein `meta`-Element innerhalb des `head`-Elements erstellen.

```js
assert.exists(document.querySelector('head > meta'));
```

Sie sollten dem `meta`-Tag ein `charset` von `UTF-8` geben.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    --fcc-editable-region--
    
    --fcc-editable-region--
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
