---
id: 6133acc353338c0bba9cb553
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Zuletzt ist im `head` das `title`-Element nützlich für Screenreader, um den Inhalt einer Seite zu verstehen. Außerdem ist es ein wichtiger Teil der _SEO_.

Geben Sie Ihrer Seite einen `title`, der beschreibend und prägnant ist.

# --hints--

Sie sollten ein `title`-Element zum `head` hinzufügen.

```js
assert.exists(document.querySelector('head > title'));
```

Sie sollten das `title` nicht länger als 60 Zeichen machen.

```js
assert.isAtMost(document.querySelector('head > title')?.textContent?.length, 60);
```

Versuchen Sie, mit Ihrem `title`-Element beschreibender zu sein. _Hinweis: Mindestens 15 Zeichen_

```js
assert.isAtLeast(document.querySelector('head > title')?.textContent?.length, 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
