---
id: 6133d11ef548f51f876149e3
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Navigation ist ein zentraler Bestandteil der Barrierefreiheit, und Screenreader sind darauf angewiesen, dass Sie die Struktur Ihrer Seite bereitstellen. Dies erreichen Sie mit semantischen HTML-Elementen.

Fügen Sie Ihrer Seite ein `header` und ein `main` Element hinzu.

Das `header` Element wird verwendet, um die Seite einzuführen und ein Navigationsmenü bereitzustellen.

Das `main` Element enthält den Hauptinhalt Ihrer Seite.

# --hints--

Sie sollten ein `header` Element zum `body` hinzufügen.

```js
assert.exists(document.querySelector('body > header'));
```

Sie sollten ein `main` Element zum `body` hinzufügen.

```js
assert.exists(document.querySelector('body > main'));
```

Das `header` Element sollte vor dem `main` Element stehen.

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    
  </body>
--fcc-editable-region--
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
