---
id: 613e2546d0594208229ada50
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Innerhalb des `header` geben Sie Kontext zur Seite, indem Sie ein `img`-, `h1`- und `nav`-Element verschachteln.

Das `img` sollte auf `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg` verweisen, ein `id` mit dem Wert `logo` haben und einen `alt`-Text von `freeCodeCamp` besitzen.

Das `h1` sollte den Text `HTML/CSS Quiz` enthalten.

# --hints--

Sie sollten dem `img`-Element ein `header`-Element hinzufügen.

```js
assert.exists(document.querySelector('header > img'));
```

Sie sollten dem `h1`-Element ein `header`-Element hinzufügen.

```js
assert.exists(document.querySelector('header > h1'));
```

Sie sollten dem `nav`-Element ein `header`-Element hinzufügen.

```js
assert.exists(document.querySelector('header > nav'));
```

Sie sollten die `img`-, `h1`- und `nav`-Elemente in dieser Reihenfolge platzieren.

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

Sie sollten dem `img`-Element ein `src`-Attribut mit dem Wert `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg` geben.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

Sie sollten dem `img`-Element ein `id`-Attribut mit dem Wert `logo` geben.

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

Sie sollten dem `img`-Element ein `alt`-Attribut mit dem Wert `freeCodeCamp` geben.

```js
assert.equal(document.querySelector('img')?.alt, 'freeCodeCamp');
```

Sie sollten dem `h1`-Element den Text `HTML/CSS Quiz` geben.

```js
assert.include(document.querySelector('h1')?.innerText.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
