---
id: 613e275749ebd008e74bb62e
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Derzeit nimmt das `img` seine Standardgröße an, die zu groß ist. CSS hat eine `max`-Funktion, die den größten Wert aus einer Menge von durch Kommas getrennten Werten zurückgibt. Zum Beispiel:

```css
img {
  width: max(250px, 25vw);
}
```

Im obigen Beispiel beträgt die Breite des Bildes 250px, wenn die Breite des Viewports weniger als 1000 Pixel beträgt. Wenn die Breite des Viewports größer als 1000 Pixel ist, beträgt die Breite des Bildes 25vw. Das liegt daran, dass 25vw 25 % der Viewport-Breite entspricht.

Skalieren Sie das Bild mit seinem `id` als Selektor und setzen Sie die `width` auf das Maximum von `10rem` oder `18vw`.

# --hints--

Sie sollten den `#logo`-Selektor verwenden, um das `img`-Element anzusprechen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

Sie sollten dem `img` eine `width` von `max(10rem, 18vw)` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
