---
id: 61408f155e798909b6908712
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Um die Navigation auf der Seite zu ermöglichen, fügen Sie eine ungeordnete Liste mit den folgenden drei Listenelementen hinzu:

- `INFO`
- `HTML`
- `CSS`

Der Text der Listenelemente sollte in Anker-Tags eingeschlossen sein.

# --hints--

Sie sollten ein `ul`-Element innerhalb des `nav` verschachteln.

```js
assert.equal(document.querySelectorAll('nav > ul')?.length, 1);
```

Sie sollten drei `li`-Elemente innerhalb des `ul`-Elements verschachteln.

```js
assert.equal(document.querySelectorAll('nav > ul > li')?.length, 3);
```

Sie sollten jeweils ein `a`-Element innerhalb jedes `li`-Elements verschachteln.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.length, 3);
```

Dem ersten `a`-Element sollten Sie den Text `INFO` geben.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')[0]?.textContent.trim(), 'INFO');
```

Dem zweiten `a`-Element sollten Sie den Text `HTML` geben.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')[1]?.textContent.trim(), 'HTML');
```

Dem dritten `a`-Element sollten Sie den Text `CSS` geben.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')[2]?.textContent.trim(), 'CSS');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}
```
