---
id: 6140827cff96e906bd38fc2b
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Wie im <a href="https://design-style-guide.freecodecamp.org/" target="_blank" rel="noopener noreferrer nofollow">freeCodeCamp Style Guide</a> beschrieben, sollte das Logo ein Seitenverhältnis von `35 / 4` beibehalten und einen Innenabstand um den Text haben.

Ändern Sie zuerst das `background-color` in `#0a0a23`, damit Sie das Logo sehen können. Verwenden Sie dann die Eigenschaft `aspect-ratio`, um das gewünschte Seitenverhältnis auf `35 / 4` einzustellen. Fügen Sie schließlich einen `padding` von `0.4rem` rundherum hinzu.

# --hints--

Sie sollten `#logo` eine `background-color` von `#0a0a23` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.backgroundColor, 'rgb(10, 10, 35)');
```

Sie sollten die Eigenschaft `aspect-ratio` verwenden.

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio);
```

Sie sollten die Eigenschaft `height` nicht verwenden.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.height);
```

Sie sollten die Eigenschaft `width` nicht ändern.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

Sie sollten `img` ein `aspect-ratio` von `35 / 4` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio, '35 / 4');
```

Sie sollten `img` einen `padding` von `0.4rem` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.padding, '0.4rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

#logo {
  width: max(10rem, 18vw);
  --fcc-editable-region--
  
  --fcc-editable-region--
}

```
