---
id: 6143931a113bb80c45546287
title: Schritt 23
challengeType: 0
dashedName: step-23
---

# --description--

Fügen Sie innerhalb des zweiten `section`-Elements zwei `div`-Elemente mit einer Klasse von `question-block` hinzu.

Dann fügen Sie innerhalb jedes `div.question-block`-Elements ein `h3`-Element mit aufsteigenden Zahlen beginnend bei `1` und ein `fieldset`-Element mit einer Klasse von `question` hinzu.

# --hints--

Sie sollten zwei `div`-Elemente innerhalb des zweiten `section`-Elements verschachteln.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div')?.length, 2);
```

Dem ersten neuen `div`-Element sollten Sie eine Klasse von `question-block` geben.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div')?.[0]?.className, 'question-block');
```

Dem zweiten neuen `div`-Element sollten Sie eine Klasse von `question-block` geben.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div')?.[1]?.className, 'question-block');
```

Sie sollten innerhalb jedes `h3`-Elements ein `div.question-block`-Element verschachteln.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > h3')?.length, 2);
```

Dem ersten `h3`-Element sollten Sie den Text `1` geben.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > h3')[0]?.textContent.trim(), '1');
```

Dem zweiten `h3`-Element sollten Sie den Text `2` geben.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > h3')[1]?.textContent.trim(), '2');
```

Sie sollten innerhalb jedes `fieldset`-Elements ein `div.question-block`-Element verschachteln.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > fieldset')?.length, 2);
```

Das erste `fieldset`-Element sollten Sie nach dem ersten `h3`-Element platzieren.

```js
assert.exists(document.querySelector('section:nth-of-type(2) > div.question-block > h3 + fieldset'));
```

Das zweite `fieldset`-Element sollten Sie nach dem zweiten `h3`-Element platzieren.

```js
assert.exists(document.querySelector('section:nth-of-type(2) > div.question-block:nth-of-type(2) > h3 + fieldset'));
```

Dem ersten `fieldset`-Element sollten Sie eine Klasse von `question` geben.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > fieldset')?.[0]?.className, 'question');
```

Dem zweiten `fieldset`-Element sollten Sie eine Klasse von `question` geben.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > fieldset')?.[1]?.className, 'question');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a href="#student-info">INFO</a></li>
          <li><a href="#html-questions">HTML</a></li>
          <li><a href="#css-questions">CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">Student Info</h2>
          <div class="info">
            <label for="student-name">Name:</label>
            <input type="text" name="student-name" id="student-name" />
          </div>
          <div class="info">
            <label for="student-email">Email:</label>
            <input type="email" name="student-email" id="student-email" />
          </div>
          <div class="info">
            <label for="birth-date">Date of Birth:</label>
            <input type="date" name="birth-date" id="birth-date" />
          </div>
        </section>
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
          --fcc-editable-region--
          
          --fcc-editable-region--
        </section>
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

h1,
h2 {
  font-family: Verdana, Tahoma;
}

h2 {
  border-bottom: 4px solid #dfdfe2;
}


```
