---
id: 614ccc21ea91ef1736b9b578
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Für diesen Workshop wurde Ihnen ein grundlegender Boilerplate-Code bereitgestellt.

Beginnen Sie diese Barrierefreiheitsreise, indem Sie Ihrem `lang`-Attribut zum `html`-Element hinzufügen. Dies hilft Screenreadern, die Sprache der Seite zu erkennen.

Vergessen Sie auch nicht, Ihr Stylesheet mit der Seite zu verknüpfen.

# --hints--

Sie sollten dem `html`-Element ein `lang`-Attribut geben. _Hinweis: Für Englisch können Sie den Wert `en` verwenden._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

Sie sollten ein `link`-Element in Ihrem `head`-Element haben.

```js
assert.exists(document.querySelector("head>link"));
```

Ihr `link`-Element sollte `rel="stylesheet"` und `href="styles.css"` enthalten.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>
    
--fcc-editable-region--
  </head>
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
