---
id: 68b94f23f7b3ee14078c89c5
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop verwenden Sie das `iframe`-Element, um ein Video anzuzeigen. Der grundlegende HTML-Boilerplate-Code wurde für Sie vorbereitet.

Beginnen Sie damit, ein `h1`-Element mit dem Text `iframe Video Display` zu erstellen.

# --hints--

Sie sollten ein `h1`-Element haben.

```js
assert.exists(document.querySelector('h1'))
```

Ihr `h1`-Element sollte den Text `iframe Video Display` enthalten.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), 'iframe Video Display')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
