---
id: 68ba9c4f1688914d72876094
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

In der ersten Lektion zum `iframe`-Element haben Sie gelernt, dass es sich um ein ersetztes Element handelt, genau wie `img`. Das bedeutet, dass es auch die `width`- und `height`-Attribute annehmen kann, um festzulegen, wie hoch und breit es sein soll.

Geben Sie Ihrem `iframe`-Element eine `width` von `560` und eine `height` von `315`.

# --hints--

Ihr `iframe`-Element sollte ein `width`-Attribut mit dem Wert `560` haben.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('width'), '560')
```

Ihr `iframe`-Element sollte ein `height`-Attribut mit dem Wert `315` haben.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('height'), '315')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>

    <iframe
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
