---
id: 68ba9c4f1688914d72876096
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Eines der Attribute ist `allow`. Es ist wie eine Berechtigungsliste, die dem Browser mitteilt, welche Funktionen der `iframe` verwenden darf.

Hier ist ein `iframe`-Element mit dem `allow`-Attribut:

```html
<iframe
  allow="accelerometer autoplay clipboard-write encrypted-media gyroscope picture-in-picture web-share"
></iframe>
```

FĂĽgen Sie das `allow`-Attribut mit den Werten `accelerometer`, `autoplay` und `clipboard-write` hinzu.

`accelerometer` erlaubt dem `iframe`, Bewegungssensoren zu verwenden, damit es Dinge wie das Kippen und Drehen des GerĂ¤ts erkennen kann. `autoplay` lĂ¤sst das Video automatisch starten, und `clipboard-write` erlaubt dem iframe, Daten in die Zwischenablage des Benutzers zu schreiben.

# --hints--

Ihr `iframe`-Element sollte ein `allow`-Attribut mit dem Wert `accelerometer` haben.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /accelerometer/)
```

Das `allow`-Attribut Ihres `iframe`-Elements sollte `autoplay` als einen seiner Werte enthalten.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /autoplay/)
```

Das `allow`-Attribut Ihres `iframe`-Elements sollte `clipboard-write` als einen seiner Werte enthalten.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /clipboard-write/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
