---
id: 68ba9c4f1688914d72876095
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Das `iframe`-Element erhält außerdem ein `src`-Attribut mit einem Wert, der die URL oder den Pfad der anzuzeigenden Ressource angibt.

Fügen Sie Ihrem `src`-Element ein `https://www.youtube.com/embed/I0_951_MPE0`-Attribut mit dem Wert `iframe` hinzu.

An dieser Stelle sollten Sie das Video auf der Seite sehen, aber es gibt noch weitere Attribute, die Sie hinzufügen müssen.

# --hints--

Ihr `iframe`-Element sollte ein `src`-Attribut mit dem Wert `https://www.youtube.com/embed/I0_951_MPE0` haben.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('src'), 'https://www.youtube.com/embed/I0_951_MPE0')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
