---
id: 68ba9c4f1688914d72876093
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Erstellen Sie jetzt ein `iframe`-Element. Fügen Sie noch nichts darin ein.

# --hints--

Sie sollten ein `iframe`-Element haben.

```js
assert.exists(document.querySelector('iframe'))
```

Ihr `iframe`-Element sollte leer sein.

```js
assert.equal(document.querySelector('iframe')?.innerHTML.trim(), '')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
