---
id: 68ba9c4f1688914d72876098
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Das nächste Attribut, das Sie hinzufügen, ist `referrerpolicy`. Es ist die Regel, die bestimmt, wie viele Details Sie teilen, wenn Ihre Seite eine Verbindung zu einer anderen Seite herstellt.

Fügen Sie das `referrerpolicy`-Attribut hinzu und setzen Sie es auf `strict-origin-when-cross-origin`. Dadurch wird die vollständige Adresse auf derselben Website geteilt, nur der Name der Website auf anderen Websites und nichts auf unsicheren Websites.

# --hints--

Ihr `iframe`-Element sollte ein `referrerpolicy`-Attribut mit dem Wert `strict-origin-when-cross-origin` haben.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('referrerpolicy'), 'strict-origin-when-cross-origin')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
      allow="accelerometer autoplay clipboard-write encrypted-media gyroscope web-share"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
