---
id: 68ba9c4f1688914d72876097
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Fügen Sie `encrypted-media`, `gyroscope` und `web-share` zu den bestehenden Werten im `allow`-Attribut hinzu.

Diese drei ermöglichen die Verwendung von verschlüsselten Media-Erweiterungen zum Schutz des Videos, gewähren Zugriff auf die Bewegungs- und Orientierungssensoren des Geräts und erlauben das Teilen des iframe-Inhalts über die nativen Freigabedialoge des Geräts.

# --hints--

Das `allow`-Attribut Ihres `iframe`-Elements sollte `encrypted-media` als einen seiner Werte enthalten.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /encrypted-media/)
```

Das `allow`-Attribut Ihres `iframe`-Elements sollte `gyroscope` als einen seiner Werte enthalten.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /gyroscope/)
```

Das `allow`-Attribut Ihres `iframe`-Elements sollte `web-share` als einen seiner Werte enthalten.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /web-share/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      allow="accelerometer autoplay clipboard-write"
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
