---
id: 69a92f0b9cc04bb0d5327bb5
title: 實作堆疊
challengeType: 26
dashedName: implement-a-stack
---

# --description--

在這個實驗中，你將使用函式實作一個堆疊資料結構。堆疊是一種後進先出（LIFO）的資料結構，元素會從頂端被加入和移除。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `initStack` 函式，傳回一個物件，該物件的 `collection` 屬性設為空陣列。`collection` 的最後一個元素代表堆疊的頂端。
1. 你應該有一個 `push` 函式，接受一個堆疊物件和一個元素作為引數，並將該元素加入堆疊頂端，也就是放在堆疊的 `collection` 陣列末端。
1. 你應該有一個 `pop` 函式，接受一個堆疊物件作為引數，移除並傳回堆疊頂端的元素，如果沒有元素則傳回 `undefined`。
1. 你應該有一個 `peek` 函式，接受一個堆疊物件作為引數，傳回堆疊頂端的元素但不移除它，如果沒有元素則傳回 `undefined`。
1. 你應該有一個 `isEmpty` 函式，接受一個堆疊物件作為引數，當堆疊沒有元素時傳回 `true`，否則傳回 `false`。
1. 你應該有一個 `clear` 函式，接受一個堆疊物件作為引數，移除堆疊 `collection` 中的所有元素。

**注意：** 大多數測試依賴 `initStack` 和 `push` 函式。請先實作它們，因為 `pop`、`peek`、`isEmpty` 和 `clear` 的測試需要先將元素加入堆疊。

# --hints--

你應該有一個 `push` 函式。

```js
assert.isFunction(push);
```

`push` 函式應該將元素加入堆疊頂端。

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

你應該有一個 `pop` 函式。

```js
assert.isFunction(pop);
```

`pop` 函式應該移除並傳回堆疊頂端的元素。

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
assert.strictEqual(pop(stack), 'third');
assert.deepEqual(stack.collection, ['first', 'second']);
assert.strictEqual(pop(stack), 'second');
assert.deepEqual(stack.collection, ['first']);
```

如果堆疊為空，`pop` 函式應該傳回 `undefined`。

```js
const stack = initStack();
assert.isUndefined(pop(stack));
```

`pop` 函式應該正確傳回假值（falsy values）。

```js
const stack = initStack();
push(stack, '');
assert.strictEqual(pop(stack), '');
assert.lengthOf(stack.collection, 0);
```

你應該有一個 `peek` 函式。

```js
assert.isFunction(peek);
```

`peek` 函式應該傳回堆疊頂端的元素但不移除它。

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.strictEqual(peek(stack), 'second');
assert.strictEqual(peek(stack), 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

如果堆疊為空，`peek` 函式應該傳回 `undefined`。

```js
const stack = initStack();
assert.isUndefined(peek(stack));
```

`peek` 函式應該正確傳回假值（falsy values）。

```js
const stack = initStack();
push(stack, 0);
assert.strictEqual(peek(stack), 0);
```

你應該有一個 `isEmpty` 函式。

```js
assert.isFunction(isEmpty);
```

當堆疊為空時，`isEmpty` 函式應該傳回 `true`。

```js
const stack = initStack();
assert.isTrue(isEmpty(stack));
```

當堆疊非空時，`isEmpty` 函式應該傳回 `false`。

```js
const stack = initStack();
push(stack, 'element');
assert.isFalse(isEmpty(stack));
```

當堆疊頂端元素為假值（falsy）時，`isEmpty` 函式應該傳回 `false`。

```js
const stack = initStack();
push(stack, false);
assert.isFalse(isEmpty(stack));
```

你應該有一個 `clear` 函式（程式）。

```js
assert.isFunction(clear);
```

`clear` 函式應該移除堆疊中的所有元素。

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
clear(stack);
assert.lengthOf(stack.collection, 0);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function initStack() {
  return {
    collection: []
  };
}

function print(stack) {
  console.log(stack.collection);
}

function push(stack, element) {
  stack.collection.push(element);
}

function pop(stack) {
  return stack.collection.pop();
}

function peek(stack) {
  return stack.collection[stack.collection.length - 1];
}

function isEmpty(stack) {
  return stack.collection.length === 0;
}

function clear(stack) {
  stack.collection = [];
}
```
