---
id: 69a92f0b9cc04bb0d5327bb5
title: 實作堆疊
challengeType: 26
dashedName: implement-a-stack
---

# --description--

在這個實驗中，你將使用函式實作一個堆疊資料結構。堆疊是一種後進先出（LIFO）的資料結構，元素會從頂端被加入和移除。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `initStack` 函式，回傳一個物件，該物件有一個 `collection` 屬性，並設為空陣列。
1. 你應該有一個 `push` 函式，能將元素加入堆疊頂端。
1. 你應該有一個 `pop` 函式，能移除並返回堆疊頂端的元素，若沒有元素則返回 `undefined`。
1. 你應該有一個 `peek` 函式，能返回堆疊頂端的元素但不移除它，若沒有元素則返回 `undefined`。
1. 你應該有一個 `isEmpty` 函式，當堆疊沒有元素時返回 `true`，否則返回 `false`。
1. 你應該有一個 `clear` 函式，能移除堆疊中所有元素。

**注意：** 大多數測試依賴 `initStack` 和 `push` 函式。請先實作它們，因為 `pop`、`peek`、`isEmpty` 和 `clear` 的測試需要先將元素加入堆疊。

# --hints--

你應該有一個 `push` 函式。

```js
assert.isFunction(push);
```

`push` 函式應該將元素加入堆疊頂端。

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

你應該有一個 `pop` 函式。

```js
assert.isFunction(pop);
```

`pop` 函式應該移除並返回堆疊頂端的元素。

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
assert.strictEqual(pop(stack), 'third');
assert.strictEqual(pop(stack), 'second');
```

`pop` 函式應該在堆疊為空時返回 `undefined`。

```js
const stack = initStack();
assert.isUndefined(pop(stack));
```

`pop` 函式應該正確返回假值。

```js
const stack = initStack();
push(stack, '');
assert.strictEqual(pop(stack), '');
assert.lengthOf(stack.collection, 0);
```

你應該有一個 `peek` 函式。

```js
assert.isFunction(peek);
```

`peek` 函式應該返回堆疊頂端的元素但不移除它。

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.strictEqual(peek(stack), 'second');
assert.strictEqual(peek(stack), 'second');
```

`peek` 函式應該在堆疊為空時返回 `undefined`。

```js
const stack = initStack();
assert.isUndefined(peek(stack));
```

`peek` 函式應該正確返回假值。

```js
const stack = initStack();
push(stack, 0);
assert.strictEqual(peek(stack), 0);
```

你應該有一個 `isEmpty` 函式。

```js
assert.isFunction(isEmpty);
```

`isEmpty` 函式應該在堆疊為空時返回 `true`。

```js
const stack = initStack();
assert.isTrue(isEmpty(stack));
```

`isEmpty` 函式應該在堆疊非空時返回 `false`。

```js
const stack = initStack();
push(stack, 'element');
assert.isFalse(isEmpty(stack));
```

`isEmpty` 函式應該在頂端元素為假值時返回 `false`。

```js
const stack = initStack();
push(stack, false);
assert.isFalse(isEmpty(stack));
```

你應該有一個 `clear` 函式。

```js
assert.isFunction(clear);
```

`clear` 函式應該移除堆疊中所有元素。

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
clear(stack);
assert.lengthOf(stack.collection, 0);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function initStack() {
  return {
    collection: []
  };
}

function print(stack) {
  console.log(stack.collection);
}

function push(stack, element) {
  stack.collection.push(element);
}

function pop(stack) {
  return stack.collection.pop();
}

function peek(stack) {
  return stack.collection[stack.collection.length - 1];
}

function isEmpty(stack) {
  return stack.collection.length === 0;
}

function clear(stack) {
  stack.collection = [];
}
```
