---
id: 587d824b367417b2b2512c48
title: Assert.isOK와 Assert.isNotOK 사용하기
challengeType: 2
forumTopicId: 301607
dashedName: use-assert-isok-and-assert-isnotok
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 시작 프로젝트를 기반으로 하고 있습니다.

`isOk()`는 참 같은 값을 테스트하고, `isNotOk()`는 거짓 같은 값을 테스트합니다.

참 같은 값과 거짓 같은 값에 대해 더 알고 싶다면, <a href="https://www.freecodecamp.org/learn/javascript-algorithms-and-data-structures/basic-algorithm-scripting/falsy-bouncer" target="_blank" rel="noopener noreferrer nofollow">Falsy Bouncer</a> 도전 과제를 시도해 보세요.

# --instructions--

`tests/1_unit-tests.js` 테스트 모음 내 `#3` 레이블이 붙은 테스트 안의 `Basic Assertions`에서, 각 `assert`를 `assert.isOk()` 또는 `assert.isNotOk()` 중 하나로 바꿔 테스트가 통과하도록 하세요 (`true`로 평가되어야 합니다). 어서션에 전달된 인수는 변경하지 마세요.

# --hints--

모든 테스트가 통과해야 합니다.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

첫 번째 어서션에 올바른 메서드를 선택하세요 - `isOk` 대 `isNotOk`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isNotOk', 'Null is falsy');
```

두 번째 어서션에 올바른 메서드를 선택하세요 - `isOk` 대 `isNotOk`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'isOk', 'A string is truthy');
```

세 번째 어서션에 올바른 메서드를 선택하세요 - `isOk` 대 `isNotOk`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[2].method, 'isOk', 'true is truthy');
```
