---
id: 587d824b367417b2b2512c49
title: 진릿값 테스트
challengeType: 2
forumTopicId: 301596
dashedName: test-for-truthiness
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 시작 프로젝트를 기반으로 구축되고 있습니다.

`isTrue()`는 불리언 값 `true`을 테스트하고 `isNotTrue()`는 불리언 값 `true`이 아닌 어떤 값이 주어져도 통과합니다.

```js
assert.isTrue(true, 'This will pass with the boolean value true');
assert.isTrue('true', 'This will NOT pass with the string value "true"');
assert.isTrue(1, 'This will NOT pass with the number value 1');
```

`isFalse()`와 `isNotFalse()`도 존재하며, 이들은 각각 true에 해당하는 값과 비슷하게 동작하지만 불리언 값 `false`을 찾습니다.

# --instructions--

`tests/1_unit-tests.js` 내 `#4` 스위트의 `Basic Assertions` 레이블이 붙은 테스트에서, 각 `assert`을 `assert.isTrue` 또는 `assert.isNotTrue`로 바꿔 테스트가 통과하도록 하세요 (평가 결과가 `true`이어야 합니다). 어서션에 전달된 인수는 변경하지 마세요.

# --hints--

모든 테스트가 통과해야 합니다.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

첫 번째 어서션에 적합한 메서드를 선택하세요 - `isTrue` 대 `isNotTrue`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isTrue', 'True is true');
```

두 번째 어서션에 적합한 메서드를 선택하세요 - `isTrue` 대 `isNotTrue`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'isTrue',
    'Double negation of a truthy value is true'
  );
```

세 번째 어서션에 적합한 메서드를 선택하세요 - `isTrue` 대 `isNotTrue`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'isNotTrue',
    'A truthy object is not true - neither is a false one'
  );
```
