---
id: 587d824c367417b2b2512c4c
title: ".deepEqual과 .notDeepEqual로 깊은 동등성 어서션하기"
challengeType: 2
forumTopicId: 301587
dashedName: assert-deep-equality-with--deepequal-and--notdeepequal
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 시작 프로젝트를 기반으로 하고 있습니다.

`deepEqual()`는 두 객체가 깊은 동등성임을 어서션합니다.

# --instructions--

`tests/1_unit-tests.js` 테스트 스위트 내 `#7` 레이블이 붙은 테스트 안 `Equality`에서 각 `assert`를 `assert.deepEqual` 또는 `assert.notDeepEqual` 중 하나로 바꿔 테스트가 통과하도록 하세요 (평가 결과가 `true`이어야 합니다). 어서션에 전달된 인수는 변경하지 마세요.

# --hints--

모든 테스트가 통과해야 합니다.

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
})();
```

첫 번째 어서션에 적합한 메서드를 선택하세요 - `deepEqual` 대 `notDeepEqual`.

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'deepEqual',
    'The order of the keys does not matter'
  );
})();
```

두 번째 어서션에 적합한 메서드를 선택하세요 - `deepEqual` 대 `notDeepEqual`.

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'notDeepEqual',
    'The position of elements within an array does matter'
  );
})();
```
