---
id: 587d824b367417b2b2512c4b
title: 엄격한 동등성을 어서션하려면 삼중 등호를 사용하세요
challengeType: 2
forumTopicId: 301610
dashedName: use-the-triple-equals-to-assert-strict-equality
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 시작 프로젝트를 기반으로 하고 있습니다.

`strictEqual()`는 `===`를 사용해 객체를 비교합니다.

# --instructions--

`tests/1_unit-tests.js` 스위트 내 `#6` 레이블이 붙은 테스트 아래 `Equality`에서 각 `assert`를 `assert.strictEqual` 또는 `assert.notStrictEqual`로 바꿔 테스트가 통과하도록 하세요 (평가 결과가 `true`이어야 합니다). 어서션에 전달된 인수는 변경하지 마세요.

# --hints--

모든 테스트가 통과해야 합니다.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

첫 번째 어서션에 적합한 메서드를 선택하세요 - `strictEqual` 대 `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'notStrictEqual',
    'with strictEqual the type must match'
  );
```

두 번째 어서션에 적합한 메서드를 선택하세요 - `strictEqual` 대 `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'strictEqual', '3*2 = 6...');
```

세 번째 어서션에 적합한 메서드를 선택하세요 - `strictEqual` 대 `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'strictEqual',
    "6 * '2' is 12. Types match !"
  );
```

네 번째 어서션에 적합한 메서드를 선택하세요 - `strictEqual` 대 `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[3].method,
    'notStrictEqual',
    'Even if they have the same elements, the Arrays are notStrictEqual'
  );
```
