---
id: 587d824c367417b2b2512c4e
title: 한 값이 다른 값보다 작거나 같음을 테스트하기
challengeType: 2
forumTopicId: 301606
dashedName: test-if-one-value-is-below-or-at-least-as-large-as-another
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 시작 프로젝트를 기반으로 하고 있습니다.

# --instructions--

`tests/1_unit-tests.js` 테스트 모음 안에 `#9` 레이블이 붙은 `Comparisons` 내에서, 각 `assert`을 `assert.isBelow` 또는 `assert.isAtLeast`로 바꿔 테스트가 통과하도록 만드세요 (평가 결과가 `true`이어야 합니다). 어서션에 전달된 인수는 변경하지 마세요.

# --hints--

모든 테스트가 통과해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

첫 번째 어서션에 적합한 메서드를 선택하세요 - `isBelow` 대 `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isAtLeast',
  '5 is at least (>=) 5'
);
```

두 번째 어서션에 적합한 메서드를 선택하세요 - `isBelow` 대 `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'isAtLeast',
  '2 * Math.random() is at least 0'
);
```

세 번째 어서션에 적합한 메서드를 선택하세요 - `isBelow` 대 `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'isBelow', '1 is smaller than 2');
```

네 번째 어서션에 적합한 메서드를 선택하세요 - `isBelow` 대 `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'isBelow',
  '2/3 (0.6666) is smaller than 1'
);
```
