---
id: 587d824d367417b2b2512c51
title: 배열에 항목이 포함되어 있는지 테스트하기
challengeType: 2
forumTopicId: 301603
dashedName: test-if-an-array-contains-an-item
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 시작 프로젝트를 기반으로 만들어지고 있습니다.

# --instructions--

`tests/1_unit-tests.js` 테스트 모음 안에 `#12` 레이블이 붙은 `Arrays` 내에서, 각 `assert`을 `assert.include` 또는 `assert.notInclude` 중 하나로 바꿔 테스트가 통과하도록 만드세요 (평가 결과가 `true`가 되어야 합니다). 어서션에 전달되는 인수는 변경하지 마세요.

# --hints--

모든 테스트가 통과해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

첫 번째 어서션에 적합한 메서드를 선택하세요 - `include` 대 `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notInclude',
  "It's summer in july..."
);
```

두 번째 어서션에 적합한 메서드를 선택하세요 - `include` 대 `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'include',
  'JavaScript is a back-end language !!'
);
```
