---
id: 587d824f367417b2b2512c5b
title: Chai-HTTP를 사용해 API 응답에 기능 테스트 실행 IV - PUT 메서드
challengeType: 2
forumTopicId: 301591
dashedName: run-functional-tests-on-an-api-response-using-chai-http-iv---put-method
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 시작 프로젝트를 기반으로 구축되고 있습니다.

이 연습은 이전 연습과 유사합니다.

이제 `PUT` 요청을 테스트하는 방법을 알았으니, 직접 처음부터 해보세요.

# --instructions--

`tests/2_functional-tests.js` 내에서 `'Send {surname: "da Verrazzano"}'` 테스트(`// #4`)를 수정하고 `put`와 `send` 메서드를 사용해 `'/travellers'` 엔드포인트를 테스트하세요.

다음 JSON 객체를 PUT 요청과 함께 보내세요:

```json
{
  "surname": "da Verrazzano"
}
```

`request.end` 콜백 내에서 다음을 확인하세요:

1. `status`가 `200`이어야 합니다.
2. `type`이 `application/json`이어야 합니다.
3. `body.name`가 `Giovanni`이어야 합니다.
4. `body.surname`가 `da Verrazzano`이어야 합니다.

위 어서션 순서를 따르세요 — 이 순서에 의존합니다. 완료 후에는 `assert.fail()`를 반드시 제거하세요.

# --hints--

모든 테스트가 통과해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

`res.status`가 200인지 테스트해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

`res.type`가 `'application/json'`인지 테스트해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.type');
assert.match(data.assertions[1].args[1], /('|")application\/json\1/);
```

`res.body.name`가 `'Giovanni'`인지 테스트해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'equal');
assert.equal(data.assertions[2].args[0], 'res.body.name');
assert.match(data.assertions[2].args[1], /('|")Giovanni\1/);
```

`res.body.surname`가 `'da Verrazzano'`인지 테스트해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'equal');
assert.equal(data.assertions[3].args[0], 'res.body.surname');
assert.match(data.assertions[3].args[1], /('|")da Verrazzano\1/);
```
