---
id: 587d824d367417b2b2512c52
title: 값이 문자열인지 테스트하기
challengeType: 2
forumTopicId: 301599
dashedName: test-if-a-value-is-a-string
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 시작 프로젝트를 기반으로 만들어지고 있습니다.

`isString` 또는 `isNotString`은 실제 값이 문자열임을 어서션합니다.

# --instructions--

`tests/1_unit-tests.js` 스위트 내 `#13` 레이블이 붙은 테스트의 `Strings` 안에서, 각 `assert`를 `assert.isString` 또는 `assert.isNotString` 중 하나로 바꿔 테스트가 통과하도록 만드세요 (`true`로 평가되어야 합니다). 어서션에 전달된 인수는 변경하지 마세요.

# --hints--

모든 테스트가 통과해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

첫 번째 어서션에 적합한 메서드를 선택하세요 - `isString` 대 `isNotString`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isNotString',
  'A float number is not a string'
);
```

두 번째 어서션에 적합한 메서드를 선택하세요 - `isString` 대 `isNotString`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'isString',
  'environment vars are strings (or undefined)'
);
```

세 번째 어서션에 적합한 메서드를 선택하세요 - `isString` 대 `isNotString`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'isString', 'A JSON is a string');
```
