---
id: 587d824f367417b2b2512c5c
title: 헤드리스 브라우저를 사용해 동작 시뮬레이션하기
challengeType: 2
dashedName: simulate-actions-using-a-headless-browser
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 시작 프로젝트를 기반으로 하고 있습니다.

다음 도전 과제들에서는 헤드리스 브라우저를 사용해 웹 페이지와의 인간 상호작용을 시뮬레이션할 것입니다.

헤드리스 브라우저는 GUI가 없는 웹 브라우저입니다. 일반 브라우저처럼 HTML, CSS, JavaScript를 렌더링하고 해석할 수 있어 웹 페이지 테스트에 특히 유용합니다.

다음 도전 과제들에서는 추가 바이너리 설치가 필요 없는 가벼운 헤드리스 브라우저인 Zombie.js를 사용할 것입니다. 하지만 더 강력한 다른 헤드리스 브라우저 옵션도 많이 있습니다.

Mocha는 실제 테스트가 실행되기 전에 일부 코드를 실행할 수 있게 해줍니다. 이는 나머지 테스트에서 사용할 데이터베이스 항목을 추가하는 등의 작업에 유용합니다.

헤드리스 브라우저를 사용할 때는 테스트를 실행하기 전에 테스트할 페이지를 **방문**해야 합니다.

`suiteSetup` 훅은 테스트 스위트 시작 시 단 한 번만 실행됩니다.

각 테스트 전, 각 테스트 후, 또는 테스트 스위트 종료 시 코드를 실행할 수 있는 여러 다른 훅 유형도 있습니다. 자세한 내용은 Mocha 문서를 참조하세요.

# --instructions--

`tests/2_functional-tests.js` 내에서 `Browser` 선언 바로 다음에, 변수의 `site` 속성에 프로젝트 URL을 추가하세요:

```js
Browser.site = 'http://0.0.0.0:3000'; // Your URL here
```

그런 다음 `'Functional Tests with Zombie.js'` 스위트 최상위 수준에서 다음 코드를 사용해 `Browser` 객체의 새 인스턴스를 생성하세요:

```js
const browser = new Browser();
```

그리고 `suiteSetup` 훅을 사용해 `browser`를 `/` 라우트로 안내하는 다음 코드를 작성하세요. **참고**: `done`는 `browser.visit`에 콜백으로 전달되므로 직접 호출하지 마세요.

```js
suiteSetup(function(done) {
  return browser.visit('/', done);
});
```

# --hints--

모든 테스트가 통과해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=4');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```
