---
id: 587d824b367417b2b2512c4a
title: 동등함을 확인하려면 이중 등호를 사용하세요
challengeType: 2
forumTopicId: 301609
dashedName: use-the-double-equals-to-assert-equality
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 시작 프로젝트를 기반으로 구축되고 있습니다.

`equal()`는 `==`를 사용해 객체를 비교합니다.

# --instructions--

`tests/1_unit-tests.js` 스위트 내 `#5` 레이블이 붙은 테스트 아래 `Equality`에서 각 `assert`를 `assert.equal` 또는 `assert.notEqual`로 바꿔 테스트가 통과하도록 하세요 (평가 결과가 `true`이어야 합니다). 어서션에 전달된 인수는 변경하지 마세요.

# --hints--

모든 테스트가 통과해야 합니다.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

첫 번째 어서션에 대해 올바른 메서드를 선택하세요 - `equal` 대 `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'equal',
    'Numbers are coerced into strings with == '
  );
```

두 번째 어서션에 대해 올바른 메서드를 선택하세요 - `equal` 대 `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'notEqual',
    ' == compares object references'
  );
```

세 번째 어서션에 대해 올바른 메서드를 선택하세요 - `equal` 대 `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'equal',
    "6 * '2' is 12 ! It should be equal to '12'"
  );
```

네 번째 어서션에 대해 올바른 메서드를 선택하세요 - `equal` 대 `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[3].method, 'notEqual', "6 + '2' is '62'...");
```
