---
id: 587d824b367417b2b2512c47
title: 변수나 함수가 정의되었는지 테스트하기
challengeType: 2
forumTopicId: 301602
dashedName: test-if-a-variable-or-function-is-defined
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 보일러플레이트 프로젝트를 기반으로 만들고 있습니다.

# --instructions--

`tests/1_unit-tests.js` 테스트 스위트 내 `#2` 레이블이 붙은 테스트의 `Basic Assertions` 안에서, 각 `assert`를 `assert.isDefined()` 또는 `assert.isUndefined()`로 바꿔 테스트가 통과하도록 만드세요 (평가 결과가 `true`이어야 합니다). 어서션에 전달된 인수는 변경하지 마세요.

# --hints--

모든 테스트가 통과해야 합니다.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

첫 번째 어서션에 대해 올바른 메서드를 선택하세요 - `isDefined` 대 `isUndefined`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'isDefined',
    'Null is not undefined'
  );
```

두 번째 어서션에 대해 올바른 메서드를 선택하세요 - `isDefined` 대 `isUndefined`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'isUndefined',
    'Undefined is undefined'
  );
```

세 번째 어서션에 대해 올바른 메서드를 선택하세요 - `isDefined` 대 `isUndefined`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'isDefined',
    'A string is not undefined'
  );
```
