---
id: 587d824e367417b2b2512c58
title: Chai-HTTP를 사용해 API 엔드포인트에서 기능 테스트 실행하기
challengeType: 2
forumTopicId: 301593
dashedName: run-functional-tests-on-api-endpoints-using-chai-http
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 시작 프로젝트를 기반으로 구축되고 있습니다.

Mocha는 `chai-http`라는 플러그인을 사용해 API 엔드포인트 호출 같은 비동기 작업을 테스트할 수 있게 해줍니다.

다음은 `chai-http`를 사용한 `'GET /hello?name=[name] => "hello [name]"'`라는 테스트 스위트 예시입니다:

```js
suite('GET /hello?name=[name] => "hello [name]"', function () {
  test('?name=John', function (done) {
    chai
      .request(server)
      .keepOpen()
      .get('/hello?name=John')
      .end(function (err, res) {
        assert.equal(res.status, 200, 'Response status should be 200');
        assert.equal(res.text, 'hello John', 'Response should be "hello John"');
        done();
      });
  });
});
```

이 테스트는 이름을 URL 쿼리 문자열(`GET`)로 포함해 서버에 `?name=John` 요청을 보냅니다. `end` 메서드의 콜백 함수에서 응답 객체(`res`)를 받고, 이 객체는 `status` 속성을 포함합니다.

첫 번째 `assert.equal`는 상태가 `200`와 같은지 확인합니다. 두 번째 `assert.equal`는 응답 문자열(`res.text`)이 `"hello John"`와 같은지 확인합니다.

또한, 테스트 콜백 함수에 `done` 매개변수가 있는 점을 주목하세요. 비동기 작업이 완료되었음을 알리려면 테스트 끝에서 인수 없이 호출하는 것이 필요합니다.

마지막으로, `keepOpen` 메서드 바로 뒤에 `request` 메서드가 있는 점을 확인하세요. 보통은 명령줄에서 테스트를 실행하거나 자동 통합 과정의 일부로 실행하며, `chai-http`가 서버를 자동으로 시작하고 중지하게 할 수 있습니다.

하지만, 프로젝트 링크를 제출할 때 실행되는 테스트는 서버가 실행 중이어야 하므로, `keepOpen` 메서드를 사용해 `chai-http`가 서버를 중지하지 못하게 해야 합니다.

# --instructions--

`tests/2_functional-tests.js` 내에서 `'Test GET /hello with no name'` 테스트(`// #1`)를 수정해 응답의 `status`와 `text`을 단언해 테스트가 통과하도록 하세요. 단언에 전달되는 인수는 변경하지 마세요.

URL 쿼리가 없어야 합니다. 이름 URL 쿼리가 없으면 엔드포인트는 `hello Guest`로 응답합니다.

# --hints--

모든 테스트가 통과해야 합니다

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

`res.status` == 200인지 테스트하세요

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

`res.text` == `'hello Guest'`인지 테스트하세요

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.text');
assert.match(data.assertions[1].args[1], /('|")hello Guest\1/);
```
