---
id: 587d824d367417b2b2512c53
title: 문자열에 부분 문자열이 포함되어 있는지 테스트하기
challengeType: 2
forumTopicId: 301597
dashedName: test-if-a-string-contains-a-substring
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 보일러플레이트 프로젝트를 기반으로 만들고 있습니다.

`include()`와 `notInclude()`은 문자열에도 사용할 수 있습니다! `include()`는 실제 문자열에 예상하는 부분 문자열이 포함되어 있음을 어서션합니다.

# --instructions--

`tests/1_unit-tests.js` 테스트 스위트 내 `#14` 레이블이 붙은 `Strings` 안에서, 각 `assert`를 `assert.include` 또는 `assert.notInclude`로 바꿔 테스트가 통과하도록 만드세요 (평가 결과가 `true`가 되어야 합니다). 어서션에 전달된 인수는 변경하지 마세요.

# --hints--

모든 테스트가 통과해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

첫 번째 어서션에 적합한 메서드를 선택하세요 - `include` 대 `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'include',
  "'Arrow' contains 'row'..."
);
```

두 번째 어서션에 적합한 메서드를 선택하세요 - `include` 대 `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'notInclude',
  "... a 'dart' doesn't contain a 'queue'"
);
```
