---
id: 587d824c367417b2b2512c4f
title: 특정 범위 내에 값이 있는지 테스트하기
challengeType: 2
forumTopicId: 301598
dashedName: test-if-a-value-falls-within-a-specific-range
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 시작 프로젝트를 기반으로 만들고 있습니다.

```javascript
.approximately(actual, expected, delta, [message])
```

`actual`가 `expected`와 +/- `delta` 범위 내에서 같은지 어서션합니다.

# --instructions--

`tests/1_unit-tests.js` 테스트 스위트 내 `#10` 레이블이 붙은 테스트의 `Comparisons` 안에서, 테스트가 통과하도록 각 `assert`를 `assert.approximately`로 변경하세요 (평가 결과가 `true`이어야 합니다).

테스트가 항상 통과하도록 최솟값(3번째 매개변수) 범위를 선택하세요. 1보다 작아야 합니다.

# --hints--

모든 테스트가 통과해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

첫 번째 어서션 `approximately(actual, expected, range)`에 올바른 범위를 선택해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'approximately');
assert.equal(
  data.assertions[0].args[2],
  0.5,
  "weirdNumbers(0.5) is in the range (0.5, 1.5]. It's within 1 +/- 0.5"
);
```

두 번째 어서션 `approximately(actual, expected, range)`에 올바른 범위를 선택해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'approximately');
assert.equal(
  data.assertions[1].args[2],
  0.8,
  "weirdNumbers(0.2) is in the range (0.2, 1.2]. It's within 1 +/- 0.8"
);
```
