---
id: 587d824f367417b2b2512c5a
title: Chai-HTTP를 사용해 API 응답에 기능 테스트 실행 III - PUT 메서드
challengeType: 2
forumTopicId: 301590
dashedName: run-functional-tests-on-an-api-response-using-chai-http-iii---put-method
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 시작 프로젝트를 기반으로 하고 있습니다.

`PUT` 요청을 테스트할 때, 종종 데이터를 함께 전송합니다. `PUT` 요청에 포함하는 데이터는 요청 본문이라고 합니다.

`PUT` 요청과 JSON 객체를 `'/travellers'` 엔드포인트로 보내려면 `chai-http` 플러그인의 `put`와 `send` 메서드를 사용할 수 있습니다:

```js
chai
  .request(server)
  .keepOpen()
  .put('/travellers')
  .send({
    "surname": [last name of a traveller of the past]
  })
  ...
```

라우트는 다음과 같이 응답합니다:

```json
{
  "name": [first name],
  "surname": [last name],
  "dates": [birth - death years]
}
```

`'/travellers'` 엔드포인트에 대한 다양한 응답은 서버 코드를 참조하세요.

# --instructions--

`tests/2_functional-tests.js` 내에서 `'Send {surname: "Colombo"}'` 테스트(`// #3`)를 수정하고 `put`과 `send` 메서드를 사용해 `'/travellers'` 엔드포인트를 테스트하세요.

PUT 요청과 함께 다음 JSON 객체를 전송하세요:

```json
{
  "surname": "Colombo"
}
```

`request.end` 콜백 내에서 다음을 확인하세요:

1. `status`가 `200`이어야 합니다
2. `type`가 `application/json`이어야 합니다
3. `body.name`가 `Cristoforo`이어야 합니다
4. `body.surname`가 `Colombo`이어야 합니다

위 어서션 순서를 따르세요 - 이 순서에 의존합니다. 완료 후에는 `assert.fail()`를 반드시 제거하세요.

# --hints--

모든 테스트가 통과해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

`res.status`가 200인지 테스트하세요.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

`res.type`가 `'application/json'`인지 테스트하세요.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.type');
assert.match(data.assertions[1].args[1], /('|")application\/json\1/);
```

`res.body.name`가 `'Cristoforo'`인지 테스트하세요.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'equal');
assert.equal(data.assertions[2].args[0], 'res.body.name');
assert.match(data.assertions[2].args[1], /('|")Cristoforo\1/);
```

`res.body.surname`가 `'Colombo'`인지 테스트하세요.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'equal');
assert.equal(data.assertions[3].args[0], 'res.body.surname');
assert.match(data.assertions[3].args[1], /('|")Colombo\1/);
```
