---
id: 587d824e367417b2b2512c55
title: 객체에 속성이 있는지 테스트하기
challengeType: 2
forumTopicId: 301604
dashedName: test-if-an-object-has-a-property
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 시작 프로젝트를 기반으로 만들어지고 있습니다.

`property`는 실제 객체가 특정 속성을 가지고 있는지 어서션합니다.

# --instructions--

`tests/1_unit-tests.js` 스위트 내 `#16` 레이블이 붙은 테스트의 `Objects` 안에서, 각 `assert`를 `assert.property` 또는 `assert.notProperty` 중 하나로 바꿔 테스트가 통과하도록 만드세요 (평가 결과가 `true`이어야 합니다). 어서션에 전달된 인수는 변경하지 마세요.

# --hints--

모든 테스트가 통과해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

첫 번째 어서션에 적합한 메서드를 선택하세요 - `property` 대 `notProperty`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notProperty',
  'A car has not wings'
);
```

두 번째 어서션에 적합한 메서드를 선택하세요 - `property` 대 `notProperty`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'property',
  'planes have engines'
);
```

세 번째 어서션에 적합한 메서드를 선택하세요 - `property` 대 `notProperty`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'property', 'Cars have wheels');
```
