---
id: 587d824a367417b2b2512c46
title: JavaScript 어서션 작동 방식 배우기
challengeType: 2
forumTopicId: 301589
dashedName: learn-how-javascript-assertions-work
---

# --description--

이 도전 과제를 진행하려면 다음 방법 중 하나를 사용해 코드를 작성하세요:

- <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">이 GitHub 저장소</a>를 복제하여 로컬에서 도전 과제를 완료하세요.
- 원하는 사이트 빌더를 사용해 프로젝트를 완료하세요. GitHub 저장소의 모든 파일을 반드시 포함하세요.

# --instructions--

`tests/1_unit-tests.js` 테스트 모음 내 `#1` 레이블이 붙은 `Basic Assertions` 안에서, 각 `assert`을 `assert.isNull` 또는 `assert.isNotNull`로 바꿔 테스트가 통과하도록 하세요 (`true`로 평가되어야 합니다). 어서션에 전달된 인수는 변경하지 마세요.

# --hints--

모든 테스트가 통과해야 합니다.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

첫 번째 어서션에 적합한 메서드를 선택하세요 - `isNull` 대 `isNotNull`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isNull', 'Null is null');
```

두 번째 어서션에 적합한 메서드를 선택하세요 - `isNull` 대 `isNotNull`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'isNotNull', '1 is not null');
```
