---
id: 587d8250367417b2b2512c5d
title: 헤드리스 브라우저를 사용해 기능 테스트 실행하기
challengeType: 2
forumTopicId: 301595
dashedName: run-functional-tests-using-a-headless-browser
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 시작 프로젝트를 기반으로 만들어지고 있습니다.

페이지에는 입력 폼이 있습니다. 이 폼은 AJAX 요청으로 `PUT /travellers` 엔드포인트에 데이터를 보냅니다.

요청이 성공적으로 완료되면, 클라이언트 코드는 응답에 포함된 정보를 담은 `<div>`를 DOM에 추가합니다.

다음은 Zombie.js를 사용해 폼과 상호작용하는 방법의 예시입니다:

```js
test('Submit the surname "Polo" in the HTML form', function (done) {
  browser.fill('surname', 'Polo').then(() => {
    browser.pressButton('submit', () => {
      browser.assert.success();
      browser.assert.text('span#name', 'Marco');
      browser.assert.text('span#surname', 'Polo');
      browser.assert.elements('span#dates', 1);
      done();
    });
  });
});
```

먼저, `fill` 객체의 `browser` 메서드가 폼의 `surname` 필드를 `'Polo'` 값으로 채웁니다. `fill`는 프로미스를 반환하므로 `then`가 체인으로 연결됩니다.

`then` 콜백 내에서 `pressButton` 객체의 `browser` 메서드를 사용해 폼의 `submit` 이벤트 리스너를 호출합니다. `pressButton` 메서드는 비동기입니다.

그다음 AJAX 요청으로부터 응답을 받으면, 다음을 확인하는 몇 가지 어서션이 이루어집니다:

1.  응답 상태가 `200`인지
2.  `<span id='name'></span>` 요소 안의 텍스트가 `'Marco'`와 일치하는지
3.  `<span id='surname'></span>` 요소 안의 텍스트가 `'Polo'`와 일치하는지
4.  `1` 요소가 `<span id='dates'></span>` 개 존재하는지

마지막으로, 비동기 테스트이기 때문에 `done` 콜백이 호출됩니다.

# --instructions--

`tests/2_functional-tests.js` 내 `'Submit the surname "Colombo" in the HTML form'` 테스트(`// #5`)에서 다음을 자동화하세요:

1.  성을 `Colombo`로 폼에 입력하세요
2.  제출 버튼을 누르세요

그리고 `pressButton` 콜백 내에서:

1.  상태가 OK `200`인지 어서션하세요
2.  `span#name` 요소 안의 텍스트가 `'Cristoforo'`인지 어서션하세요
3.  `span#surname` 요소 안의 텍스트가 `'Colombo'`인지 어서션하세요
4.  `span#dates` 요소가 존재하고 개수가 `1`인지 어서션하세요

`assert.fail()` 호출을 반드시 제거하세요.

# --hints--

모든 테스트가 통과해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

헤드리스 브라우저 요청이 성공했는지 어서션해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'browser.success');
```

`span#name` 요소 안의 텍스트가 `'Cristoforo'`인지 어서션해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'browser.text');
assert.match(data.assertions[1].args[0], /('|")span#name\1/);
assert.match(data.assertions[1].args[1], /('|")Cristoforo\1/);
```

`span#surname` 요소 안의 텍스트가 `'Colombo'`인지 어서션해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'browser.text');
assert.match(data.assertions[2].args[0], /('|")span#surname\1/);
assert.match(data.assertions[2].args[1], /('|")Colombo\1/);
```

`span#dates` 요소가 존재하고 개수가 1인지 어서션해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'browser.elements');
assert.match(data.assertions[3].args[0], /('|")span#dates\1/);
assert.equal(data.assertions[3].args[1], 1);
```
