---
id: 587d824f367417b2b2512c59
title: Chai-HTTP II를 사용해 API 엔드포인트에서 기능 테스트 실행하기
challengeType: 2
forumTopicId: 301592
dashedName: run-functional-tests-on-api-endpoints-using-chai-http-ii
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 보일러플레이트 프로젝트를 기반으로 구축되고 있습니다.

# --instructions--

`tests/2_functional-tests.js` 내에서 `'Test GET /hello with your name'` 테스트(`// #2`)를 수정해 응답의 `status`과 `text`를 검증하여 테스트가 통과하도록 하세요.

이름을 URL 쿼리로 보내려면 `?name=<your_name>`를 라우트에 덧붙이세요. 엔드포인트는 `'hello <your_name>'`로 응답합니다.

# --hints--

모든 테스트가 통과해야 합니다

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=1');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

`res.status` == 200인지 테스트해야 합니다

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=1');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

`res.text` == `'hello <your_name>'`인지 테스트해야 합니다

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=1');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.text');
assert.match(data.assertions[1].args[1], /hello [\w\d_-]/);
```
