---
id: 587d824e367417b2b2512c56
title: 특정 자료 구조 유형인지 값 테스트하기
challengeType: 2
forumTopicId: 301601
dashedName: test-if-a-value-is-of-a-specific-data-structure-type
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 보일러플레이트 프로젝트를 기반으로 만들고 있습니다.

`#typeOf`는 `Object.prototype.toString`에 의해 결정된 주어진 문자열과 값의 유형이 일치하는지 어서션합니다.

# --instructions--

`tests/1_unit-tests.js` 스위트 내 `#17` 레이블이 붙은 테스트의 `Objects` 안에서, 각 `assert`를 `assert.typeOf` 또는 `assert.notTypeOf`로 바꿔 테스트가 통과하도록 하세요 (평가 결과가 `true`이어야 합니다). 어서션에 전달된 인수는 변경하지 마세요.

# --hints--

모든 테스트가 통과해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

첫 번째 어서션에 적합한 메서드를 선택하세요 - `typeOf` 대 `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'typeOf',
  'myCar is typeOf Object'
);
```

두 번째 어서션에 적합한 메서드를 선택하세요 - `typeOf` 대 `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'typeOf',
  'Car.model is a String'
);
```

세 번째 어서션에 적합한 메서드를 선택하세요 - `typeOf` 대 `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[2].method,
  'notTypeOf',
  'Plane.wings is a Number (not a String)'
);
```

네 번째 어서션에 적합한 메서드를 선택하세요 - `typeOf` 대 `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'typeOf',
  'Plane.engines is an Array'
);
```

다섯 번째 어서션에 적합한 메서드를 선택하세요 - `typeOf` 대 `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[4].method,
  'typeOf',
  'Car.wheels is a Number'
);
```
