---
id: 5f8884f4c46685731aabfc41
title: 헤드리스 브라우저를 사용한 기능 테스트 실행 II
challengeType: 2
forumTopicId: 301594
dashedName: run-functional-tests-using-a-headless-browser-ii
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 보일러플레이트 프로젝트를 기반으로 구축되고 있습니다.

# --instructions--

`tests/2_functional-tests.js` 내에서, `'Submit the surname "Vespucci" in the HTML form'` 테스트(`// #6`)를 자동화하세요:

1. 성을 `Vespucci`로 입력하세요
2. 제출 버튼을 누르세요

그리고 `pressButton` 콜백 내에서:

1. 상태가 OK인지 확인하세요 `200`
2. `span#name` 요소 안의 텍스트가 `'Amerigo'`인지 확인하세요
3. `span#surname` 요소 안의 텍스트가 `'Vespucci'`인지 확인하세요
4. `span#dates` 요소가 존재하고 그 개수가 `1`인지 확인하세요

`assert.fail()` 호출을 제거하는 것을 잊지 마세요.

# --hints--

모든 테스트가 통과해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

헤드리스 브라우저 요청이 성공했는지 확인해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'browser.success');
```

`span#name` 요소 안의 텍스트가 `'Amerigo'`인지 확인해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'browser.text');
assert.match(data.assertions[1].args[0], /('|")span#name\1/);
assert.match(data.assertions[1].args[1], /('|")Amerigo\1/);
```

`span#surname` 요소 안의 텍스트가 `'Vespucci'`인지 확인해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'browser.text');
assert.match(data.assertions[2].args[0], /('|")span#surname\1/);
assert.match(data.assertions[2].args[1], /('|")Vespucci\1/);
```

`span#dates` 요소가 존재하고 그 개수가 1인지 확인해야 합니다.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'browser.elements');
assert.match(data.assertions[3].args[0], /('|")span#dates\1/);
assert.equal(data.assertions[3].args[1], 1);
```
