---
id: 63ee3f71381756f9716727ef
title: CSS 기초 연습 A
challengeType: 25
dashedName: css-foundations-exercise-a
---

# --description--

**목표:**  
이 연습에서는 외부 CSS, 내부 CSS, 인라인 CSS 세 가지 방법을 모두 사용해 HTML 파일에 CSS를 추가하는 연습을 합니다. 외부 및 내부 방법으로 스타일을 추가할 때는 타입 선택자만 사용하세요. 색상은 `RGB` 또는 `HEX` 값 대신 키워드(예: "blue")를 사용하세요.

## 사용자 스토리

- `div` 요소에 텍스트가 보여야 합니다.  
  - `red` 배경, `white` 텍스트, `32px` 폰트 크기, 텍스트는 가운데 정렬되고 `bold`이어야 합니다.  
  - `div` 요소의 CSS는 외부에서 타입 선택자를 사용해 추가해야 합니다.
- `p` 요소에 텍스트가 보여야 합니다.  
  - `green` 배경, `white` 텍스트, `18px` 폰트 크기여야 합니다.  
  - `p` 요소의 CSS는 내부에서 타입 선택자를 사용해 추가해야 합니다.
- `button` 요소에 텍스트가 보여야 합니다.  
  - `button` 요소는 `orange` 배경과 `18px` 폰트 크기를 가져야 합니다.  
  - `button` 요소의 CSS는 인라인 스타일로 추가해야 합니다.

# --hints--

텍스트가 포함된 `div` 요소가 하나 있어야 합니다.

```js
const divElement = document.querySelector('div');

assert.isNotNull(divElement);
assert.isAtLeast(divElement?.innerText.length, 1);
```

`div` 요소 스타일을 포함하는 외부 스타일시트가 있어야 합니다.

```js
const styleSheet = new __helpers.CSSHelp(document).getStyleSheet();
const isExternal = styleSheet?.ownerNode.classList.contains('fcc-injected-styles');
const divStyle = new __helpers.CSSHelp(document).getStyle('div');

assert.isTrue(isExternal);
assert.isNotNull(divStyle);
```

`div` 요소에는 내부 또는 인라인 스타일로 CSS가 추가되지 않아야 합니다.

```js
const styleElement = document.querySelector('style:not([class])');

assert.isNotTrue(styleElement?.innerText.includes('div'));
assert.isNotTrue(document.querySelector('div')?.hasAttribute('style'));
```

`div` 요소는 `background-color`가 `red`이고 `color`가 `white`이어야 합니다.

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const divBGColor = divStyle?.getPropertyValue('background-color');
const divColor = divStyle?.getPropertyValue('color');

assert.equal(divBGColor, 'red');
assert.equal(divColor, 'white');
```

`div` 요소는 `font-weight`이 `bold`, `font-size`이 `32px`, `text-align`이 `center`로 설정되어야 합니다.

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const textAlign = divStyle?.getPropertyValue('text-align');
const fontSize = divStyle?.getPropertyValue('font-size');
const fontWeight = divStyle?.getPropertyValue('font-weight');

assert.equal(textAlign, 'center');
assert.equal(fontSize, '32px');
assert.equal(fontWeight,'bold');
```

텍스트가 포함된 `p` 요소가 하나 있어야 합니다.

```js
const pElement = document.querySelector('p');

assert.isNotNull(pElement);
assert.isAtLeast(pElement?.innerText.length, 1)
```

`p` 요소는 `style` 요소를 사용해 내부적으로 스타일이 추가되어야 합니다.

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let isStyled = false;

if (rules && rules.selectorText === 'p') {
  isStyled = true;
}

assert.isTrue(isStyled);
```

`p` 요소는 `font-size`가 `18px`이고 `color`가 `white`로 설정되어야 합니다.

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let fontSize, color;

if (rules && rules.selectorText === 'p') {
  fontSize = rules.style.fontSize;
  color = rules.style.color;
}

assert.equal(fontSize, "18px");
assert.equal(color, 'white');
```

텍스트가 포함된 `button` 요소가 하나 있어야 합니다.

```js
const btnElement = document.querySelector('button');

assert.isNotNull(btnElement);
assert.isAtLeast(btnElement?.innerText.length, 1);
```

`button` 요소는 인라인 스타일을 가져야 합니다.

```js
assert.isTrue(document.querySelector('button')?.hasAttribute('style'));
```
 
`button` 요소는 `background-color`가 `orange`로 설정되어야 합니다.

```js
assert.equal(document.querySelector('button')?.style.backgroundColor, 'orange')
```

`button` 요소는 `font-size`가 `18px`로 설정되어야 합니다.

```js
assert.equal(document.querySelector('button')?.style.fontSize, '18px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Exercise A</title>
  </head>
  <body>

  </body>
</html>
```

```css
/* styles.css */
```

# --solutions--

```html
<!DOCTYPE html>
<html>
<head>
  <title>My Styling Example</title>
  <link rel="stylesheet" type="text/css" href="styles.css">
  <style>
    p {
      background-color: green;
      color: white;
      font-size: 18px;
    }
  </style>
</head>
<body>
  <div>Hello World!</div>
  <p>This is a paragraph.</p>
  <button style="background-color: orange; font-size: 18px;">Click Me</button>
</body>
</html>
```

```css
div {
  background-color: red;
  color: white;
  font-size: 32px;
  text-align: center;
  font-weight: bold;
}
```
