---
id: 63ee3fe9381756f9716727f1
title: CSS 기초 연습 C
challengeType: 25
dashedName: css-foundations-exercise-c
---

# --description--

이전 도전 과제에서 한 요소에 여러 클래스를 추가해 두 가지 다른 규칙을 적용한 것에서 조금 더 발전해 봅시다.

1. 첫 번째 `button` 요소에는 검은색 배경과 흰색 텍스트가 보여야 합니다.
1. 두 번째 `button` 요소에는 노란색 배경이 보여야 합니다.
1. 선택자 목록을 사용해 `28px` 크기의 폰트 크기를 설정해야 합니다.
1. 선택자 목록에 `Helvetica`와 `Times New Roman`를 포함하고 `sans-serif`를 대체 수단으로 하는 폰트 목록이 있어야 합니다.
1. 각 요소에는 고유한 클래스 이름이 있어야 합니다.
1. 두 요소가 공유하는 스타일에 대한 선택자 목록이 있어야 합니다.

# --hints--

첫 번째 요소에는 `black` 배경이 있어야 합니다.

```js
const classes = document.querySelectorAll('button')?.[0].classList;

const style = new __helpers.CSSHelp(document).getStyle(`.${classes[0]}`);

assert.equal(style?.backgroundColor, 'black');

```

첫 번째 요소에는 `white` 텍스트 색상이 있어야 합니다.

```js

const classes = document.querySelectorAll('button')?.[0].classList;

const style = new __helpers.CSSHelp(document).getStyle(`.${classes[0]}`);

assert.equal(style?.color, 'white');

```

선택자 목록을 사용해 두 요소 모두에 `font-size`를 `28px`로 설정해야 합니다.

```js
const classOne = document.querySelectorAll('button')?.[0].classList?.[0];
const classTwo = document.querySelectorAll('button')?.[1].classList?.[0];

function eitherOr() {
  const a = new __helpers.CSSHelp(document)
  return a.getStyle(`.${classOne}, .${classTwo}`) ?? a.getStyle(`.${classTwo}, .${classOne}`);
}
assert.equal(eitherOr()?.fontSize, '28px');
```

각 요소에는 고유한 클래스 이름이 있어야 합니다.

```js
const elementOneClasses = document.querySelectorAll('button')?.[0].classList;
const elementTwoClasses = document.querySelectorAll('button')?.[1].classList;

for(let i = 0; i < elementOneClasses.length; i++){
    assert(![...elementTwoClasses].includes(elementOneClasses[i]));
}

```

두 요소가 공유하는 스타일에 대한 선택자 목록이 있어야 합니다.

```js
const classOne = document.querySelectorAll('button')?.[0].classList?.[0];
const classTwo = document.querySelectorAll('button')?.[1].classList?.[0];

function eitherOr() {
  const a = new __helpers.CSSHelp(document)
  return a.getStyle(`.${classOne}, .${classTwo}`) ?? a.getStyle(`.${classTwo}, .${classOne}`);
}
assert.exists(eitherOr());
```

선택자 목록에 `Helvetica`와 `Times New Roman`를 포함하고 `sans-serif`를 대체 수단으로 하는 폰트 목록이 있어야 합니다.

```js
const classOne = document.querySelectorAll('button')?.[0].classList?.[0];
const classTwo = document.querySelectorAll('button')?.[1].classList?.[0];

function eitherOr() {
  const a = new __helpers.CSSHelp(document)
  return a.getStyle(`.${classOne}, .${classTwo}`) ?? a.getStyle(`.${classTwo}, .${classOne}`);
}

assert.equal(eitherOr()?.fontFamily, 'Helvetica, "Times New Roman", sans-serif');
```

각 요소에는 하나의 클래스만 있어야 합니다.

```js
const elementOneClasses = document.querySelectorAll('button')?.[0].classList;
const elementTwoClasses = document.querySelectorAll('button')?.[1].classList;

assert(elementOneClasses.length === 1 && elementTwoClasses.length === 1);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Grouping Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <button>Click Me!</button>
    <button>No, Click Me!</button>
  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
  <html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Grouping Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <button class="inverted">Click Me!</button>
    <button class="fancy">No, Click Me!</button>
  </body>
</html>
```

```css
.inverted,
.fancy {
  font-family: Helvetica, "Times New Roman", sans-serif;
  font-size: 28px;
}

.inverted {
  background-color: black;
  color: white;
}

.fancy {
  background-color: yellow;
}
```
