---
id: 63ee3fe4381756f9716727f0
title: CSS 기초 연습 B
challengeType: 25
dashedName: css-foundations-exercise-b
---

# --description--

**목표:** 제공된 HTML 파일에는 여러 요소가 있으며, 여기에 클래스 또는 ID 속성을 추가해야 합니다. 그런 다음 올바른 선택자 구문을 사용해 제공된 CSS 파일에 규칙을 추가해야 합니다.

## 사용자 스토리

1. 목록의 홀수 번호 요소 모두에 `yellow` 배경이 보여야 합니다.

1. 목록의 홀수 번호 요소 모두에 `class` 선택자가 사용되어야 합니다.

1. 목록의 두 번째 요소는 `blue` 텍스트와 `font-size`가 `36px`여야 합니다.

1. 두 번째 요소의 `font-size`와 텍스트 색상은 `id` 속성을 사용해 설정해야 합니다.

1. 목록의 세 번째 요소는 `font-size`가 `24px`여야 합니다.

1. 세 번째 요소의 `font-size`는 `class` 속성을 사용해 설정해야 합니다.

1. 목록의 네 번째 요소는 `red` 배경, `font-size`가 `24px`, 그리고 `font-weight`가 `bold`여야 합니다.

1. 네 번째 요소의 `font-size`는 `class` 속성으로 설정해야 합니다. `font-weight`과 색상은 `id` 속성으로 설정해야 합니다.

# --hints--

모든 홀수 요소는 `class` 속성을 가져야 합니다.

```js
const p = Array.from(document.querySelectorAll('P'));

const everyPHasClass = p?.every((paragraph) => paragraph.classList.length > 0);

assert(everyPHasClass);
```

홀수 요소는 `background-color`가 `yellow`여야 합니다.

```js
const p = Array.from(document.querySelectorAll('P'));

const everyPHasBackgroundColor = p?.every((paragraph) => {
  const style = getComputedStyle(paragraph);
  
  return style?.backgroundColor === 'rgb(255, 255, 0)';
})

assert.equal(everyPHasBackgroundColor, true);
```

두 번째 요소는 `blue` 텍스트와 `font-size`가 `36px`여야 합니다.

```js
const secondElementId = document.querySelectorAll('div')?.[0]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${secondElementId}`);

assert.equal(style?.color, 'blue');
assert.equal(style?.fontSize, '36px');
```

세 번째 요소는 텍스트와 `font-size`가 `24px`여야 합니다.

```js
const thirdElement = document.querySelectorAll('p')?.[1];

assert(thirdElement?.innerText?.length > 0);

const thirdElementClasses = Array.from(thirdElement?.classList?.values());

assert(thirdElementClasses.some(thirdElementClass => {

  const style = new __helpers.CSSHelp(document).getStyle(`.${thirdElementClass}`);

  return style?.fontSize === '24px';

}))

``` 

네 번째 요소는 `font-size`가 `24px`여야 합니다.

```js
const fourthElementClass = document.querySelectorAll('div')?.[1]?.classList[0];

const style = new __helpers.CSSHelp(document).getStyle(`.${fourthElementClass}`);

assert(style?.fontSize === '24px');
```

네 번째 요소는 `red` `background-color`를 가져야 합니다.

```js
const fourthElement = document.querySelectorAll('div')?.[1]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${fourthElement}`);

assert(style?.backgroundColor === 'red');
```

네 번째 요소는 `font-weight`가 `bold`여야 합니다.

```js
const fourthElement = document.querySelectorAll('div')?.[1]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${fourthElement}`);

assert(style?.fontWeight === 'bold');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Class and ID Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <p>Number 1 - I'm a class!</p>
    <div>Number 2 - I'm one ID.</div>
    <p>Number 3 - I'm a class, but cooler!</p>
    <div>Number 4 - I'm another ID.</div>
    <p>Number 5 - I'm a class!</p>
  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Class and ID Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <p class="odd">Number 1 - I'm a class!</p>
    <div id="two">Number 2 - I'm one ID.</div>
    <p class="odd adjust-font-size">Number 3 - I'm a class, but cooler!</p>
    <div id="four" class="adjust-font-size">Number 4 - I'm another ID.</div>
    <p class="odd">Number 5 - I'm a class!</p>
  </body>
</html>
```

```css
.odd {
  background-color: yellow;
  font-family: Verdana, "DejaVu Sans", sans-serif;
}

.adjust-font-size {
  font-size: 24px;
}

#two {
  color: blue;
  font-size: 36px;
}

#four {
  background-color: red;
  font-weight: bold;
}
```
