---
id: 69008059d43d6751385883c1
title: Налагодження помилок операторів інкременту та декременту у проблемному застосунку
challengeType: 26
dashedName: debug-increment-and-decrement-operator-errors-in-a-buggy-app
---

# --description--

Ваш друг попросив допомогти йому налагодити свій код. Він практикується з операторами інкременту та декременту, але отримує несподівані результати у консолі.

Ваше завдання — допомогти йому виявити помилки та оновити код так, щоб вихідні дані в консолі відповідали очікуваним результатам.

Ви можете звернутися до попередніх уроків про оператори інкременту та декременту, а також про префіксну та постфіксну нотації.

**Історія користувача:**

1. У вас має бути змінна з ідентифікатором `updatedCounter`, яка правильно дає значення `11` за допомогою оператора інкременту. Оновіть вираз, щоб використовувати правильну префіксну або постфіксну нотацію замість жорстко заданого значення.
2. У вас має бути змінна з ідентифікатором `finalScore`, яка правильно дає значення `8` за допомогою оператора інкременту. Оновіть вираз, щоб використовувати правильну префіксну або постфіксну нотацію замість жорстко заданого значення.
3. У вас має бути змінна з ідентифікатором `updatedCoins`, яка правильно дає значення `2` за допомогою оператора декременту. Оновіть вираз, щоб використовувати правильну префіксну або постфіксну нотацію замість жорстко заданого значення.
4. У вас має бути змінна з ідентифікатором `newHealth`, яка правильно дає значення `7` за допомогою оператора декременту. Оновіть вираз, щоб використовувати правильну префіксну або постфіксну нотацію замість жорстко заданого значення.

# --hints--

У вас має бути змінна з ідентифікатором `updatedCounter`, яка використовує префіксну нотацію інкременту (`++counter`).

```js
assert.match(code, /updatedCounter\s*=\s*\+\+counter\s*;/);
assert.strictEqual(typeof updatedCounter, 'number');
```

`updatedCounter` має дорівнювати `11`.

```js
assert.strictEqual(updatedCounter, 11);
```

У вас має бути змінна з ідентифікатором `finalScore`, яка використовує постфіксну нотацію інкременту (`score++`).

```js
assert.match(code, /finalScore\s*=\s*score\+\+\s*;/);
assert.strictEqual(typeof finalScore, 'number');
```

`finalScore` має дорівнювати `8`.

```js
assert.strictEqual(finalScore, 8);
```

У вас має бути змінна з ідентифікатором `updatedCoins`, яка використовує префіксну нотацію декременту (`--coins`).

```js
assert.match(code, /updatedCoins\s*=\s*--coins\s*;/);
assert.strictEqual(typeof updatedCoins, 'number');
```

`updatedCoins` має дорівнювати `2`.

```js
assert.strictEqual(updatedCoins, 2);
```

У вас має бути змінна з ідентифікатором `newHealth`, яка використовує постфіксну нотацію декременту (`health--`).

```js
assert.match(code, /newHealth\s*=\s*health--\s*;/);
assert.strictEqual(typeof newHealth, 'number');
```

`newHealth` має дорівнювати `7`.

```js
assert.strictEqual(newHealth, 7);
```

# --seed--

## --seed-contents--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = counter++;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = ++score;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = coins--;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = --health;
console.log(`Health is still ${newHealth}`);
```

# --solutions--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = ++counter;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = score++;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = --coins;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = health--;
console.log(`Health is still ${newHealth}`);
```
