---
id: 60f0286404aefb0562a4fdf9
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Add a `title` and `meta` element inside the `head` element. Give your project a title of `Registration Form`, and add the `charset` attribute with a value of `utf-8` to your `meta` element.

# --hints--

Tu código debe tener un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Tu elemento `title` debe estar dentro del elemento `head`.

```js
assert.exists(document.querySelector('head > title'));
```

El título de tu proyecto debe ser `Registration Form`. Remember, the casing and spelling matters for the title.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

Debes crear un elemento `meta` dentro del elemento `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

Debes dar al elemento `meta` un atributo `charset` con el valor `utf-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

Tu elemento `meta` debe ser un elemento vacío, no tiene una etiqueta final `</meta>`.

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
