---
id: 60fab4a123ce4b04526b082b
title: Paso 28
challengeType: 0
dashedName: step-28
---

# --description--

Por ahora, los usuarios pueden enviar el formulario sin seleccionar los inputs de tipo radio. Aunque anteriormente usaste el atributo `required` para indicar que un input es obligatorio, en este caso, añadir `required` a ambos inputs no funcionará, ya que transmitirá información incorrecta a los usuarios.

Para resolver esto, puedes proporcionar contexto de lo que se necesita, añadiendo un elemento `legend` (leyenda) con el texto `Account type (required)` antes de los elementos `label` dentro del segundo `fieldset`. Después añade un atributo `checked` al input con el texto `Personal` para asegurarse de que el formulario se envía con los datos requeridos.

# --hints--

Tus elementos `input` deben permanecer con un `type` con el valor `radio`.

```js
assert.equal(document.querySelectorAll('fieldset:nth-child(2) input[type="radio"]')?.length, 2);
```

Debes añadir un elemento `legend` dentro del segundo elemento `fieldset`.

```js
assert.equal(document.querySelectorAll('fieldset:nth-of-type(2) > legend')?.length, 1);
```

Debes añadir el texto `Account type (required)` al elemento `legend`.

```js
assert.equal(document.querySelector('fieldset:nth-of-type(2) > legend')?.textContent, 'Account type (required)');
```

Debes añadir un atributo `checked` al input con el texto `Personal`.

```js
assert.isTrue(document.querySelector('fieldset:nth-child(2) input[type="radio"]')?.hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset>

        <label><input type="radio" name="account-type" /> Personal</label>
        <label><input type="radio" name="account-type" /> Business</label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
