---
id: 60fab8367d35de04e5cb7929
title: Paso 31
challengeType: 0
dashedName: step-31
---

# --description--

Add the text `I accept the terms and conditions` immediately after the `input` element in the newly added `label`. Then link the text `terms and conditions` to the following location:

```md
https://www.freecodecamp.org/news/terms-of-service/
```

# --hints--

Debes añadir el texto `I accept the terms and conditions` a la `label` que sigue al tercer `fieldset`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label')?.innerText.trim(), 'I accept the terms and conditions');
```

Debes usar un elemento `a` para convertir en un enlace a terms and conditions.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label a'));
```

Debes añadir el nuevo texto inmediatamente después del `input` en el elemento `label`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input + a'));
```

Debes dar al elemento `a` un atributo `href` con el valor `https://www.freecodecamp.org/news/terms-of-service/`.

```js
assert.match(document.querySelector('fieldset:nth-child(3) + label > input + a')?.href, /https:\/\/www\.freecodecamp\.org\/news\/terms-of-service\/?/);
```

Debes usar `a` para solo envolver el texto `terms and conditions`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input + a')?.textContent.trim(), 'terms and conditions');
```

El texto dentro de tu elemento anchor tiene espacio en blanco extra inicial o final. Los únicos espacios en el texto del elemento anchor deben estar entre las palabras `terms`, `and` y `conditions`.

```js
const nestedAnchor = document.querySelector('fieldset:nth-child(3) + label > input + a')?.textContent;
const innerContent = nestedAnchor.innerHTML;
assert.isNotTrue(/^\s+|\s+$/.test(innerContent));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--
      <label for="terms-and-conditions"><input id="terms-and-conditions" type="checkbox" required /></label>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
