---
id: 62ff8e998d3e7eae14d6ae3b
title: Paso 30
challengeType: 0
dashedName: step-30
---

# --description--

Debe confirmar que el usuario ha leído los términos y condiciones (terms and conditions).

Añade un elemento `label`. Dentro del nuevo elemento `label`, agrega un elemento `input` con un atributo `type` con el valor `checkbox`. Haz que este elemento `input` sea requerido (`required`) para que los usuarios no puedan registrarse sin leer los términos y condiciones.

Añade un atributo `id` y un atributo `for` y dales el valor `terms-and-conditions` para la accesibilidad.

# --hints--

Debes añadir un elemento `label` después del tercer elemento `fieldset`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label'));
```

Debes añadir un `input` al elemento `label`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input'));
```

Debes añadir un atributo `type` con el valor `checkbox` al elemento `input`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.type, 'checkbox');
```

Debes agregar un atributo `required` al elemento `input`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.required, true);
```

El elemento `input` debe tener un `id` con el valor `terms-and-conditions`.

```js
assert(document.querySelector('fieldset:nth-child(3) + label > input')?.matches('#terms-and-conditions'))
```

El elemento `label` debe tener un atributo `for` con el valor `terms-and-conditions`.

```js
assert(document.querySelector('fieldset:nth-child(3) + label')?.matches('label[for="terms-and-conditions"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--

--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
