---
id: 6140cbeec34e970dfe75e710
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Створіть селектор для вашого `.wheel` елемента. Почніть із встановлення `border` у `2px solid black`, `border-radius` у `50%` та `margin-left` у `50px`.

# --hints--

У вас має бути `.wheel` селектор.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel'));
```

Ваш `.wheel` селектор має мати властивість `border`, встановлену в `2px solid black`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.border === '2px solid black');
```

Ваш `.wheel` селектор має мати властивість `border-radius`, встановлену в `50%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.borderRadius === '50%');
```

Ваш `.wheel` селектор має мати властивість `margin-left`, встановлену в `50px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.marginLeft === '50px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
