---
id: 6140c7e645d8e905819f1dd4
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Почніть зі стандартного шаблонного коду. Додайте своє оголошення `DOCTYPE`, свій елемент `html` з мовою, встановленою на англійську, елементи `head` та `body`.

Додайте свій елемент `meta` для правильного `charset`, елемент `title` і елемент `link` для файлу `./styles.css`.

Встановіть `title` у `Ferris Wheel`.

# --hints--

Ваш код має містити оголошення `DOCTYPE`.

```js
assert.match(code, /<!DOCTYPE/gi);
```

Після оголошення `DOCTYPE` має бути пробіл.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

Визначте тип документа як `html`.

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

Закрийте оголошення `DOCTYPE` кінцевим тегом `>` після типу.

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

Ваш елемент `html` має мати початковий тег з атрибутом `lang` зі значенням `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Ваш елемент `html` має мати кінцевий тег.

```js
assert.match(code, /<\/html\s*>/);
```

Оголошення `DOCTYPE` має бути на початку вашого HTML.

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

У вас має бути початковий тег `head`.

```js
assert.match(code, /<head\s*>/i);
```

У вас має бути кінцевий тег `head`.

```js
assert.match(code, /<\/head\s*>/i);
```

У вас має бути початковий тег `body`.

```js
assert.match(code, /<body\s*>/i);
```

У вас має бути кінцевий тег `body`.

```js
assert.match(code, /<\/body\s*>/i);
```

Ваш елемент `body` має бути першим елементом після елемента `head`.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

Ваші елементи `head` і `body` мають бути братськими.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

Ваш елемент `head` має бути першим елементом всередині елемента `html`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

Ваш елемент `head` має бути всередині елемента `html`.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Ваш елемент `body` має бути всередині елемента `html`.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Ваш код має містити елемент `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Ваш елемент `meta` має мати атрибут `charset` зі значенням `UTF-8`.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

Ваш елемент `meta` має бути всередині елемента `head`.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

Ваш код має містити елемент `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

У вашого проєкту має бути заголовок `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

Пам’ятайте, що регістр і написання мають значення для заголовка.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

Ваш код має містити елемент `link`.

```js
assert.match(code, /<link/);
```

У вас має бути один елемент `link`.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

Ваш елемент `link` має бути всередині вашого елемента `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Ваш елемент `link` має мати атрибут `rel` зі значенням `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Ваш елемент `link` має мати атрибут `href` зі значенням `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
