---
id: 6140d10d50636e14695013b2
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Створіть селектор, щоб націлитися на другий `.line` елемент. Встановіть властивість `transform` у значення `rotate(60deg)`.

Пам’ятайте, що властивість `transform` дозволяє маніпулювати формою елемента. У цьому випадку використання значення `rotate(60deg)` поверне елемент навколо його `transform-origin` точки на 60 градусів за годинниковою стрілкою.

# --hints--

Ви маєте створити `.line:nth-of-type(2)` селектор.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)'));
```

Ваш `.line:nth-of-type(2)` селектор повинен мати властивість `transform`, встановлену у значення `rotate(60deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)')?.transform === 'rotate(60deg)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

--fcc-editable-region--

--fcc-editable-region--
```
