---
id: 60a3e3396c7b40068ad6996b
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

У межах елемента `head` додайте тег `meta`, який встановлює `charset` у `UTF-8`, а також елемент `title` зі значенням `Rothko Painting`.

У межах елемента `body` додайте елемент `img` з ідентифікатором `src`, рівним `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.
# --hints--

Ваш код має містити тег `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Тег `meta` має встановлювати `charset` у `UTF-8`.

```js
assert(document.querySelector('meta').getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Ваш код має містити елемент `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

`title` має бути `Rothko Painting`.

```js
assert(document.querySelector('title').innerText === 'Rothko Painting');
```

Ваш код має містити елемент `img`.

```js
assert(document.querySelectorAll('img').length === 1);
```

Елемент `img` має мати ідентифікатор `src`, рівний `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
