---
id: 60a3e3396c7b40068ad6996a
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

До цього моменту ви маєте бути знайомі з базовими елементами, які повинна містити HTML-сторінка.

Налаштуйте свій код з оголошенням `DOCTYPE`, елементом `html` з мовою, встановленою на англійську, елементом `head` та елементом `body`.

# --hints--

Ваш код має містити оголошення `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

У вас має бути відкриваючий тег `<html>` з атрибутом `lang`, значення якого `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ваш код має містити елемент `head` всередині елемента `html`.

```js
assert(document.querySelectorAll('head').length === 1);
```

Ваш код має містити елемент `body` всередині елемента `html`.

```js
assert(document.querySelectorAll('body').length === 1);
```

Ваш елемент `head` має йти перед елементом `body`.

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
