---
id: 60a3e3396c7b40068ad6996f
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Додайте елемент `div` у `body`.

Встановіть атрибут `class` зі значенням `canvas`.

Це слугуватиме полотном для вашого малюнка.

# --hints--

Ваш код має містити елемент `div`.

```js
assert(document.querySelectorAll('div').length === 1)
```

Елемент `div` повинен мати атрибут `class` зі значенням `canvas`.

```js
assert(document.querySelector('div').className.split(' ').includes('canvas'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
