---
id: 60a3e3396c7b40068ad6996c
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

У моделі CSS-блоку кожен HTML-елемент розглядається як блок із чотирма областями.

Уявіть, що ви отримали коробку від улюбленого інтернет-магазину — вмістом є товар у коробці, або в нашому випадку, заголовок, абзац чи елемент зображення.

Змініть атрибут `src` у `<img>` з `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` на `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`.

# --hints--

Елемент `img` повинен мати ідентифікатор `src` зі значенням `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```
