---
id: 60a3e3396c7b40068ad69970
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Перш ніж почати стилізувати `div`, який ви додали, потрібно зв’язати свій CSS з HTML.

Додайте `link` елемент, щоб підключити свій `styles.css` файл. Встановіть `href` на `styles.css`, і не забудьте встановити атрибут `rel` зі значенням `stylesheet`.

# --hints--

Ваш код має містити `link` елемент.

```js
assert(/<link/.test(code))
```

Ваш `link` елемент повинен мати атрибут `rel` зі значенням `stylesheet`.

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

Ваш `link` елемент повинен мати атрибут `href` зі значенням `styles.css`.

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
