---
id: 60a3e3396c7b40068ad6996e
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Поле — це область поза межами блоку, яку можна використовувати для контролю простору між іншими блоками або елементами.

Тут нижній елемент має більше верхнє поле, що відсуває його далі вниз по сторінці.

Тепер, коли ви розумієте модель блоку CSS, давайте почнемо працювати над картиною Ротко.

Видаліть елемент `<img>`.

# --hints--

У вашому коді не має бути елемента `img`.

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
