---
id: 61438b5b66d76a6264430f2a
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Dentro de cada uno de sus nuevos elementos `a`, agregue un elemento `i` y asígneles las siguientes clases:

- Su primer elemento `i` debe tener la class `fab fa-facebook-f`
- Tu segundo elemento `i` debe tener la clase `fab fa-twitter`
- Tu tercer elemento `i` debe tener la clase `fab fa-instagram`
- Tu cuarto elemento `i` debe tener la clase `fab fa-linkedin-in`
- Tu quinto elemento `i` debe tener la clase `fab fa-youtube`

# --hints--

Debes tener cinco elementos `i`.

```js
assert.lengthOf(document.querySelectorAll('i'), 5);
```

Cada elemento `a` debe tener solo un elemento `i`.

```js
const aaaaa = document.querySelectorAll('.social-icons a');
assert.lengthOf(aaaaa, 5);
aaaaa?.forEach(
  a => {
    assert.lengthOf(a?.children, 1);
    assert.equal(a?.children?.[0]?.localName, 'i');
  }
);
```

Cada elemento `i` debe tener un atributo `class` que incluya `fab`.

```js
const iiiii = document.querySelectorAll('i');
assert.lengthOf(iiiii, 5);
iiiii?.forEach(
  i => {
    assert.isTrue(i?.classList?.contains('fab'));
  }
);
```

El primer elemento `i` debe tener un atributo `class` que incluya `fa-facebook-f`.

```js
assert.isTrue(document.querySelectorAll('i')?.[0]?.classList?.contains('fa-facebook-f'));
```

El segundo elemento `i` debe tener un atributo `class` que incluya `fa-twitter`.

```js
assert.isTrue(document.querySelectorAll('i')?.[1]?.classList?.contains('fa-twitter'));
```

El tercer elemento `i` debe tener un atributo `class` que incluya `fa-instagram`.

```js
assert.isTrue(document.querySelectorAll('i')?.[2]?.classList?.contains('fa-instagram'));
```

El cuarto elemento `i` debe tener un atributo `class` que incluya `fa-linkedin-in`.

```js
assert.isTrue(document.querySelectorAll('i')?.[3]?.classList?.contains('fa-linkedin-in'));
```

El quinto elemento `i` debe tener un atributo `class` que incluya `fa-youtube`.

```js
assert.isTrue(document.querySelectorAll('i')?.[4]?.classList?.contains('fa-youtube'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--
        <div class="social-icons">
          <a href="https://www.facebook.com/freecodecamp/">
          </a>
          <a href="https://twitter.com/freecodecamp/">
          </a>
          <a href="https://instagram.com/freecodecamp">
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
          </a>
          <a href="https://www.youtube.com/freecodecamp">
          </a>
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
