---
id: 614387cbefeeba5f3654a291
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Después de su elemento `header`, cree un `div` con la `class` establecida en `author`.

Dentro de ese `div`, crea un elemento `p` con la `class` establecida en `author-name` y dale el texto `By freeCodeCamp`. Envuelve la parte de `freeCodeCamp` en un elemento `a` con el `href` establecido en `https://freecodecamp.org`, y el `target` establecido en `_blank`.


Debajo de eso, agrega un segundo elemento `p` con la clase `publish-date` y el texto `March 7, 2019`.

# --hints--

Debes crear un nuevo elemento `div`.

```js
assert.exists(document.querySelector('div'));
```

Tu elemento `div` debe estar después de tu elemento `header`.

```js
assert.equal(document.querySelector('div')?.previousElementSibling?.localName, 'header');
```

Tu elemento `div` debe tener la `class` establecida en `author`.

```js
assert.equal(document.querySelector('div')?.className, 'author');
```

Debes crear dos nuevos elementos `p`.

```js
assert.lengthOf(document.querySelectorAll('p'), 3);
```

Tus dos nuevos elementos `p` deben estar dentro de tu elemento `div`.

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

Tu primer nuevo elemento `p` debe tener una `class` establecida en `author-name`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.className, 'author-name');
```

Tu primer nuevo elemento `p` debe tener el texto `By freeCodeCamp`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.innerText, 'By freeCodeCamp');
```

Tu segundo nuevo elemento `p` debe tener una `class` establecida en `publish-date`.

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className, 'publish-date');
```

Tu segundo nuevo elemento `p` debe tener el texto `March 7, 2019`.

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText, 'March 7, 2019');
```

Debes crear un nuevo elemento `a`.

```js
assert.exists(document.querySelector('a'));
```

Tu elemento `a` debe estar dentro de tu primer nuevo elemento `p`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName, 'a');
```

Tu elemento `a` debería tener el `href` establecido en `https://freecodecamp.org`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href'), 'https://freecodecamp.org');
```

Tu elemento `a` debe tener su `target` establecido en `_blank`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target'), '_blank');
```

Su elemento `a` debe rodear el texto `freeCodeCamp`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.innerText, 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
