---
id: 6143862a5e54455d262c212e
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro de tu `body`, crea un elemento `main`. Luego, en ese elemento, cree un `section` con un `class` establecido en `heading`.

# --hints--

Debes tener un elemento `main`.

```js
assert.exists(document.querySelector('main'));
```

Tu elemento `main` debe estar dentro de tu elemento `body`.

```js
assert.equal(document.querySelector('main')?.parentElement?.localName, 'body');
```

Debes tener un elemento `section`.

```js
assert.exists(document.querySelector('section'));
```

Tu elemento `section` debe estar dentro de tu elemento `main`.

```js
assert.equal(document.querySelector('section')?.parentElement?.localName, 'main');
```

Tu elemento `section` debe tener la `class` establecida en `heading`.

```js
assert.isTrue(document.querySelector('section')?.classList?.contains('heading'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
