---
id: 6165d3b702a5d92ad970b30c
title: Pasos 4
challengeType: 0
dashedName: step-4
---

# --description--

Después de tu elemento `img`, añade un elemento `h1` con la `class` establecida en `hero-title` y el texto establecido en `OUR NEW CURRICULUM`, seguido de un elemento `p` con la `class` establecida en `hero-subtitle` y el texto establecido en `Nuestros esfuerzos para reestructurar nuestro plan de estudios con un enfoque más basado en proyectos`.

# --hints--

Debes crear un elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Tu elemento `h1` debe estar después de tu elemento `img`.

```js
assert.equal(document.querySelector('h1')?.previousElementSibling?.localName, 'img');
```

Tu elemento `h1` debe tener la `class` establecida en `hero-title`.

```js
assert.equal(document.querySelector('h1')?.className, 'hero-title');
```

Tu elemento `h1` debe tener el texto establecido en `OUR NEW CURRICULUM`.

```js
assert.equal(document.querySelector('h1')?.textContent, 'OUR NEW CURRICULUM');
```

Debes crear un nuevo elemento `p`.

```js
assert.exists(document.querySelector('p'));
```

Tu elemento `p` debe estar después de tu elemento `h1`.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.localName, 'h1');
```

Tu elemento `p` debe tener la `class` establecida en `hero-subtitle`.

```js
assert.equal(document.querySelector('p')?.className, 'hero-subtitle');
```

Tu elemento `p` debe tener el texto establecido en `Nuestros esfuerzos para reestructurar nuestro plan de estudios con un enfoque más basado en proyectos`.

```js
assert.equal(document.querySelector('p')?.textContent?.trim()?.replace(/\s{2,}/, ' '), 'Our efforts to restructure our curriculum with a more project-based focus');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
        </header>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
