---
id: 587d781c367417b2b2512ac2
title: 여러 제목 요소에 폰트 크기 설정
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPpQNT3'
forumTopicId: 301067
dashedName: set-the-font-size-for-multiple-heading-elements
---

# --description--

`font-size` 속성은 특정 요소 내 텍스트 크기를 지정하는 데 사용됩니다. 이 규칙은 여러 요소에 적용해 페이지 내 텍스트의 시각적 일관성을 만들 수 있습니다. 이 도전 과제에서는 `h1`부터 `h6` 태그까지의 값을 설정해 제목 크기를 균형 있게 조절할 것입니다.

# --instructions--

`style` 태그 내에서 다음 태그들의 `font-size`를 설정하세요:

- `h1` 태그는 68px로 설정하세요.
- `h2` 태그는 52px로 설정하세요.
- `h3` 태그는 40px로 설정하세요.
- `h4` 태그는 32px로 설정하세요.
- `h5` 태그는 21px로 설정하세요.
- `h6` 태그는 14px로 설정하세요.

# --hints--

`font-size` 태그에 대해 `h1` 속성을 68픽셀로 설정해야 합니다.

```js
 const fontSizeOfh1 = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh1 ,'68px');
```

`font-size` 태그에 대해 `h2` 속성을 52픽셀로 설정해야 합니다.

```js
 const fontSizeOfh2 = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh2 ,'52px');
```

`font-size` 태그에 대해 `h3` 속성을 40픽셀로 설정해야 합니다.

```js
 const fontSizeOfh3 = new __helpers.CSSHelp(document).getStyle('h3')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh3 ,'40px');
```

`font-size` 태그에 대해 `h4` 속성을 32픽셀로 설정해야 합니다.

```js
 const fontSizeOfh4 = new __helpers.CSSHelp(document).getStyle('h4')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh4 , '32px');
```

`font-size` 태그에 대해 `h5` 속성을 21픽셀로 설정해야 합니다.

```js
 const fontSizeOfh5 = new __helpers.CSSHelp(document).getStyle('h5')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh5 ,'21px');
```

`font-size` 태그에 대해 `h6` 속성을 14픽셀로 설정해야 합니다.

```js
 const fontSizeOfh6 = new __helpers.CSSHelp(document).getStyle('h6')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh6 , '14px');
```

# --seed--

## --seed-contents--

```html
<style>






</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```

# --solutions--

```html
<style>
  h1 {
    font-size: 68px;
  }
  h2 {
    font-size: 52px;
  }
  h3 {
    font-size: 40px;
  }
  h4 {
    font-size: 32px;
  }
  h5 {
    font-size: 21px;
  }
  h6 {
    font-size: 14px;
  }
</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```
