---
id: 587d78a3367417b2b2512acf
title: z-index 속성으로 겹치는 요소의 위치 변경하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM94aHk'
forumTopicId: 301046
dashedName: change-the-position-of-overlapping-elements-with-the-z-index-property
---

# --description--

요소들이 겹치도록 위치가 지정되었을 때(즉, `position: absolute | relative | fixed | sticky` 사용 시), 기본적으로 HTML 마크업에서 나중에 오는 요소가 다른 요소들 위에 나타납니다. 하지만 `z-index` 속성은 요소들이 서로 쌓이는 순서를 지정할 수 있습니다. 이 값은 정수여야 하며(소수가 아닌 온전한 숫자), 요소의 `z-index` 속성 값이 높을수록 낮은 값보다 더 위쪽에 쌓입니다.

# --instructions--

클래스 이름이 `z-index`인 요소(빨간 사각형)에 `first` 속성을 추가하고 값을 2로 설정하여 다른 요소(파란 사각형)를 덮도록 하세요.

# --hints--

클래스가 `first`인 요소는 `z-index` 값이 2여야 합니다.

```js
const firstElement = document.querySelector('.first');
const firstStyle = window.getComputedStyle(firstElement); 
assert.equal(firstStyle?.zIndex, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 60%;
    height: 200px;
    margin-top: 20px;
  }

  .first {
    background-color: red;
    position: absolute;

  }
  .second {
    background-color: blue;
    position: absolute;
    left: 40px;
    top: 50px;
    z-index: 1;
  }
</style>

<div class="first"></div>
<div class="second"></div>
```

# --solutions--

```html
<style>
  div {
    width: 60%;
    height: 200px;
    margin-top: 20px;
  }

  .first {
    background-color: red;
    position: absolute;
    z-index: 2;
  }
  .second {
    background-color: blue;
    position: absolute;
    left: 40px;
    top: 50px;
    z-index: 1;
  }
</style>
<div class="first"></div>
<div class="second"></div>
```
