---
id: 587d78a3367417b2b2512ad0
title: margin 속성을 사용해 요소를 가로로 가운데 정렬하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLJqU4'
forumTopicId: 301043
dashedName: center-an-element-horizontally-using-the-margin-property
---

# --description--

또 다른 위치 지정 기법은 블록 요소를 가로로 가운데 정렬하는 것입니다. 이를 하는 한 가지 방법은 `margin`를 auto 값으로 설정하는 것입니다.

이 방법은 이미지에도 적용됩니다. 이미지는 기본적으로 인라인 요소이지만 `display` 속성을 `block`로 설정하면 블록 요소로 변경할 수 있습니다.

# --instructions--

`div`를 페이지에서 가운데 정렬하려면 `margin` 속성을 `auto` 값으로 추가하세요.

# --hints--

`div`에 `margin`를 `auto`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('div')?.margin ,'auto');
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    background-color: blue;
    height: 100px;
    width: 100px;

  }
</style>
<div></div>
```

# --solutions--

```html
<style>
  div {
    background-color: blue;
    height: 100px;
    width: 100px;
    margin: auto;
  }
</style>
<div></div>
```
