---
id: 587d78a5367417b2b2512ad7
title: 줄무늬 요소를 만들기 위해 CSS 선형 그라디언트 사용하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/c6bmQh2'
forumTopicId: 301072
dashedName: use-a-css-linear-gradient-to-create-a-striped-element
---

# --description--

`repeating-linear-gradient()` 함수는 `linear-gradient()`과 매우 비슷하지만, 주요 차이점은 지정한 그라디언트 패턴을 반복한다는 점입니다. `repeating-linear-gradient()`는 다양한 값을 받을 수 있지만, 이 도전 과제에서는 각도 값과 색상 정지 값을 사용해 작업할 것입니다.

각도 값은 그라디언트의 방향을 의미합니다. 색상 정지는 전환이 일어나는 위치를 표시하는 너비 값과 같으며, 백분율이나 픽셀 단위로 지정합니다.

코드 편집기에서 보여준 예시에서, 그라디언트는 0픽셀에서 `yellow` 색상으로 시작해 시작점에서 40픽셀 떨어진 곳에서 두 번째 색상 `blue`로 섞입니다. 다음 색상 정지도 40픽셀에 있으므로, 그라디언트는 즉시 세 번째 색상 `green`로 바뀌고, 이 색상은 그라디언트 시작점에서 80픽셀 떨어진 네 번째 색상 `red`로 섞입니다.

이 예시에서는 색상 정지를 두 개씩 짝지어 생각하면 도움이 됩니다. 두 색상이 함께 섞이는 방식입니다.

```css
0px [yellow -- blend -- blue] 40px [green -- blend -- red] 80px
```

만약 두 색상 정지 값이 같은 색상이라면, 섞임이 눈에 띄지 않고 같은 색상 사이에서만 일어나며, 다음 색상으로는 급격한 전환이 발생해 줄무늬가 만들어집니다.

# --instructions--

`repeating-linear-gradient()`을 변경해 그라디언트 각도를 `45deg`로 설정하고, 첫 두 색상 정지를 `yellow`로, 두 번째 두 색상 정지를 `black`로 설정해 줄무늬를 만드세요.

# --hints--

`repeating-linear-gradient()`의 각도는 45deg여야 합니다.

```js
assert.match(code,/background:\s*?repeating-linear-gradient\(\s*?45deg/gi);
```

`repeating-linear-gradient()`의 각도는 더 이상 90deg가 아니어야 합니다.

```js
assert.notMatch(code, /90deg/gi);
```

0픽셀에서의 색상 정지는 `yellow`이어야 합니다.

```js
assert.match(code, /yellow\s+?0(px)?/gi);
```

40픽셀에서 첫 번째 색상 정지는 `yellow`이어야 합니다.

```js
assert.match(code, /yellow\s+?40px/gi);
```

40픽셀에서 두 번째 색상 정지는 `black`이어야 합니다.

```js
assert.match(code, /yellow\s+?40px,\s*?black\s+?40px/gi);
```

80픽셀에서 마지막 색상 정지는 `black`이어야 합니다.

```js
assert.match(code, /black\s+?80px/gi);
```

# --seed--

## --seed-contents--

```html
<style>

  div{
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin:  50 auto;
    background: repeating-linear-gradient(
      90deg,
      yellow 0px,
      blue 40px,
      green 40px,
      red 80px
    );
  }

</style>

<div></div>
```

# --solutions--

```html
<style>
  div{
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin:  50 auto;
    background: repeating-linear-gradient(
      45deg,
      yellow 0px,
      yellow 40px,
      black 40px,
      black 80px
    );
  }
</style>
<div></div>
```
