---
id: 587d78a8367417b2b2512ae7
title: 키워드로 애니메이션 타이밍 변경
challengeType: 0
videoUrl: 'https://scrimba.com/c/cJKvwCM'
forumTopicId: 301045
dashedName: change-animation-timing-with-keywords
---

# --description--

CSS 애니메이션에서 `animation-timing-function` 속성은 애니메이션이 진행되는 동안 애니메이션 요소가 얼마나 빠르게 변하는지 제어합니다. 애니메이션이 주어진 시간(당신의 `animation-duration`)에 점 A에서 점 B로 움직이는 자동차라고 하면, `animation-timing-function`는 운전하는 동안 자동차가 어떻게 가속하고 감속하는지를 나타냅니다.

인기 있는 옵션에 대해 미리 정의된 여러 키워드가 있습니다. 예를 들어 기본값은 `ease`로, 천천히 시작해서 중간에 빨라졌다가 끝에 다시 느려집니다. 다른 옵션으로는 처음에 빠르게 시작해서 느려지는 `ease-out`, 처음에 느리다가 끝에 빨라지는 `ease-in`, 그리고 애니메이션 내내 일정한 속도를 적용하는 `linear`가 있습니다.

# --instructions--

id가 `ball1`과 `ball2`인 요소에 각각 `animation-timing-function` 속성을 추가하고, `#ball1`을 `linear`로, `#ball2`를 `ease-out`로 설정하세요. 두 요소가 애니메이션 중 움직이는 방식의 차이를 확인하세요. 하지만 두 요소는 `animation-duration`가 2초로 같기 때문에 함께 끝납니다.

# --hints--

id가 `animation-timing-function`인 요소의 `ball1` 속성 값은 `linear`이어야 합니다.

```js
const ballOne =document.querySelector('#ball1'); 
const ballOneStyle = window.getComputedStyle(ballOne); 

const ball1Animation = __helpers.removeWhiteSpace(ballOneStyle?.animationTimingFunction);
assert.isTrue(ball1Animation == 'linear' || ball1Animation == 'cubic-bezier(0,0,1,1)');
```

id가 `animation-timing-function`인 요소의 `ball2` 속성 값은 `ease-out`이어야 합니다.

```js
const ballTwo = document.querySelector('#ball2'); 
const ballTwoStyle = window.getComputedStyle(ballTwo); 

const ball2Animation = __helpers.removeWhiteSpace(ballTwoStyle?.animationTimingFunction);
assert.isTrue(ball2Animation == 'ease-out' || ball2Animation == 'cubic-bezier(0,0,0.58,1)');
```

# --seed--

## --seed-contents--

```html
<style>

  .balls {
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left:27%;

  }
  #ball2 {
    left:56%;

  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }

</style>

<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```

# --solutions--

```html
<style>
  .balls {
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left:27%;
    animation-timing-function: linear;
  }
  #ball2 {
    left:56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```
