---
id: 587d78a9367417b2b2512ae8
title: 베지어 곡선 작동 방식 배우기
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bDrs8'
forumTopicId: 301058
dashedName: learn-how-bezier-curves-work
---

# --description--

마지막 도전 과제에서는 `animation-timing-function` 속성과 애니메이션 지속 시간 동안 속도를 변경하는 몇 가지 키워드를 소개했습니다. CSS는 키워드 외에도 애니메이션이 어떻게 진행되는지 더 세밀하게 제어할 수 있는 베지어 곡선을 제공합니다.

CSS 애니메이션에서 베지어 곡선은 `cubic-bezier` 함수와 함께 사용됩니다. 곡선의 모양은 애니메이션이 어떻게 진행되는지를 나타냅니다. 이 곡선은 1x1 좌표계 위에 존재합니다. 이 좌표계의 X축은 애니메이션의 지속 시간(시간 척도라고 생각하세요), Y축은 애니메이션의 변화량입니다.

`cubic-bezier` 함수는 이 1x1 격자 위에 위치한 네 개의 주요 점으로 구성됩니다: `p0`, `p1`, `p2`, 그리고 `p3`. `p0`과 `p3`은 항상 각각 원점(0, 0)과 (1, 1)에 위치하는 시작점과 끝점으로 설정되어 있습니다. 나머지 두 점의 x와 y 값을 설정하며, 이 점들을 격자 내 어디에 배치하느냐에 따라 애니메이션이 따르는 곡선의 모양이 결정됩니다. CSS에서는 `p1`와 `p2` "앵커" 점의 x와 y 값을 `(x1, y1, x2, y2)` 형식으로 선언하여 이를 지정합니다. 이를 모두 합치면, 다음은 CSS 코드에서 베지어 곡선의 예시입니다:

```css
animation-timing-function: cubic-bezier(0.25, 0.25, 0.75, 0.75);
```

위 예시에서 각 점의 x와 y 값이 동일합니다(x1 = 0.25 = y1, x2 = 0.75 = y2). 기하학 수업에서 기억한다면, 이는 원점에서 (1, 1) 점까지 이어지는 직선을 의미합니다. 이 애니메이션은 애니메이션 길이 동안 요소가 선형으로 변화하는 것으로, `linear` 키워드를 사용하는 것과 같습니다. 즉, 일정한 속도로 변화합니다.

# --instructions--

id가 `ball1`인 요소에 대해 `animation-timing-function` 속성의 값을 `linear`에서 위 예시의 동등한 `cubic-bezier` 함수 값으로 변경하세요. 예시에서 주어진 점 값을 사용하세요.

# --hints--

id가 `animation-timing-function`인 요소의 `ball1` 속성 값은 선형에 해당하는 `cubic-bezier` 함수여야 합니다.

```js
const ballOne = document.querySelector('#ball1');
const ballOneStyle = window.getComputedStyle(ballOne); 
assert.equal(
  ballOneStyle?.animationTimingFunction,
    'cubic-bezier(0.25, 0.25, 0.75, 0.75)'
);
```

id가 `animation-timing-function`인 요소의 `ball2` 속성 값은 변경하지 않아야 합니다.

```js
const ballTwoElement = document.querySelector('#ball2');
const ballTwoStyle = window.getComputedStyle(ballTwoElement); 
const ball2Animation = __helpers.removeWhiteSpace(ballTwoStyle?.animationTimingFunction);
assert.isTrue(
  ball2Animation == 'ease-out' || ball2Animation == 'cubic-bezier(0,0,0.58,1)'
);
```

# --seed--

## --seed-contents--

```html
<style>

  .balls{
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left: 27%;
    animation-timing-function: linear;
  }
  #ball2 {
    left: 56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }

</style>

<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```

# --solutions--

```html
<style>

  .balls{
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left: 27%;
    animation-timing-function: cubic-bezier(0.25, 0.25, 0.75, 0.75);
  }
  #ball2 {
    left: 56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```
