---
id: 587d78a7367417b2b2512ae2
title: 왼쪽에서 오른쪽으로 요소를 페이드 아웃하여 시각적 방향 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJqqAE'
forumTopicId: 301054
dashedName: create-visual-direction-by-fading-an-element-from-left-to-right
---

# --description--

이 도전 과제에서는 애니메이션 요소의 `opacity`를 변경하여 화면 오른쪽에 도달할 때 점차 페이드 아웃되도록 만드세요.

표시된 애니메이션에서 그라디언트 배경의 둥근 요소는 `@keyframes` 규칙에 따라 애니메이션의 50% 지점까지 오른쪽으로 이동합니다.

# --instructions--

id가 `ball`인 요소를 선택하고 `opacity`에서 `50%` 속성을 0.1로 설정하여 요소가 오른쪽으로 이동할 때 페이드 아웃되도록 하세요.

# --hints--

페이드 아웃을 위한 `keyframes` 규칙은 50% 지점에서 `opacity` 속성을 0.1로 설정해야 합니다.

```js
assert.match(code,/@keyframes fade\s*?{\s*?50%\s*?{\s*?(?:left:\s*?60%;\s*?opacity:\s*?0?\.1;|opacity:\s*?0?\.1;\s*?left:\s*?60%;)/gi
);
```

# --seed--

## --seed-contents--

```html
<style>

  #ball {
    width: 70px;
    height: 70px;
    margin: 50px auto;
    position: fixed;
    left: 20%;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: fade;
    animation-duration: 3s;
  }

  @keyframes fade {
    50% {
      left: 60%;

    }
  }

</style>

<div id="ball"></div>
```

# --solutions--

```html
<style>
  #ball {
    width: 70px;
    height: 70px;
    margin: 50px auto;
    position: fixed;
    left: 20%;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: fade;
    animation-duration: 3s;
  }

  @keyframes fade {
    50% {
      left: 60%;
      opacity: 0.1;
    }
  }
</style>
<div id="ball"></div>
```
