---
id: 587d78a9367417b2b2512aea
title: 베지어 곡선을 사용해 움직임을 더 자연스럽게 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/c/c7akWUv'
forumTopicId: 301063
dashedName: make-motion-more-natural-using-a-bezier-curve
---

# --description--

이 도전 과제는 공이 저글링되는 움직임을 재현하기 위해 요소를 애니메이션합니다. 이전 도전 과제에서는 `linear`와 `ease-out` 3차 베지어 곡선을 다뤘지만, 어느 것도 저글링 움직임을 정확히 표현하지 못했습니다. 이 경우에는 베지어 곡선을 직접 맞춤 설정해야 합니다.

`animation-timing-function`는 `animation-iteration-count`가 무한대로 설정되면 각 키프레임마다 자동으로 반복됩니다. 애니메이션 지속 시간 중간에 키프레임 규칙이 설정되어 있기 때문에(`50%` 위치), 공의 위아래 움직임에서 두 개의 동일한 애니메이션 진행이 발생합니다.

다음 3차 베지어 곡선은 저글링 움직임을 시뮬레이션합니다:

```css
cubic-bezier(0.3, 0.4, 0.5, 1.6);
```

y2 값이 1보다 크다는 점에 주목하세요. 3차 베지어 곡선은 1x1 좌표계에 매핑되며 x 값은 0에서 1 사이만 허용하지만, y 값은 1보다 큰 숫자로 설정할 수 있습니다. 이로 인해 튀는 움직임이 만들어져 저글링 공을 시뮬레이션하는 데 이상적입니다.

# --instructions--

id가 `animation-timing-function`인 요소의 `green` 속성 값을 x1, y1, x2, y2가 각각 0.311, 0.441, 0.444, 1.649로 설정된 `cubic-bezier` 함수로 변경하세요.

# --hints--

id가 `animation-timing-function`인 요소의 `green` 속성 값은 지정된 x1, y1, x2, y2 값으로 된 `cubic-bezier` 함수여야 합니다.

```js
const greenElement = document.querySelector('#green');
 const greenStyle = window.getComputedStyle(greenElement);
assert.equal(greenStyle?.animationTimingFunction, 'cubic-bezier(0.311, 0.441, 0.444, 1.649)');
```

# --seed--

## --seed-contents--

```html
<style>
  .balls {
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    top: 60%;
    animation-name: jump;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 25%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 50%;
    animation-timing-function: ease-out;
  }
  #green {
    background: green;
    left: 75%;
    animation-timing-function: cubic-bezier(0.69, 0.1, 1, 0.1);
  }

  @keyframes jump {
    50% {
      top: 10%;
    }
  }
</style>
<div class="balls" id="red"></div>
<div class="balls" id="blue"></div>
<div class="balls" id="green"></div>
```

# --solutions--

```html
<style>
  .balls {
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    top: 60%;
    animation-name: jump;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 25%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 50%;
    animation-timing-function: ease-out;
  }
  #green {
    background: green;
    left: 75%;
    animation-timing-function: cubic-bezier(0.311, 0.441, 0.444, 1.649);
  }

  @keyframes jump {
    50% {
      top: 10%;
    }
  }
</style>
<div class="balls" id="red"></div>
<div class="balls" id="blue"></div>
<div class="balls" id="green"></div>
```
