---
id: 587d781c367417b2b2512ac3
title: 여러 제목 요소에 폰트 두께 설정
challengeType: 0
videoUrl: 'https://scrimba.com/c/crVWRHq'
forumTopicId: 301069
dashedName: set-the-font-weight-for-multiple-heading-elements
---

# --description--

지난 도전 과제에서 각 제목 태그의 `font-size`를 설정했는데, 이번에는 `font-weight`를 조정할 것입니다.

`font-weight` 속성은 텍스트 섹션에서 글자의 두께를 두껍거나 얇게 설정합니다.

# --instructions--

<ul><li><code>h1</code> 태그의 <code>font-weight</code>를 800으로 설정하세요.</li><li><code>h2</code> 태그의 <code>font-weight</code>를 600으로 설정하세요.</li><li><code>h3</code> 태그의 <code>font-weight</code>를 500으로 설정하세요.</li><li><code>h4</code> 태그의 <code>font-weight</code>를 400으로 설정하세요.</li><li><code>h5</code> 태그의 <code>font-weight</code>를 300으로 설정하세요.</li><li><code>h6</code> 태그의 <code>font-weight</code>를 200으로 설정하세요.</li></ul>

# --hints--

코드는 `font-weight` 태그의 `h1` 속성을 800으로 설정해야 합니다.

```js
const h1Element =document.querySelector('h1')
const h1Style = window.getComputedStyle(h1Element);
assert.equal(h1Style?.fontWeight, '800');
```

코드는 `font-weight` 태그의 `h2` 속성을 600으로 설정해야 합니다.

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.fontWeight, '600');
```

코드는 `font-weight` 태그의 `h3` 속성을 500으로 설정해야 합니다.

```js
const h3Element =document.querySelector('h3')
const h3Style = window.getComputedStyle(h3Element);
assert.equal(h3Style?.fontWeight, '500');
```

코드는 `font-weight` 태그의 `h4` 속성을 400으로 설정해야 합니다.

```js
const h4Element =document.querySelector('h4')
const h4Style = window.getComputedStyle(h4Element);
assert.equal(h4Style?.fontWeight, '400');
```

코드는 `font-weight` 태그의 `h5` 속성을 300으로 설정해야 합니다.

```js
const h5Element =document.querySelector('h5')
const h5Style = window.getComputedStyle(h5Element);
assert.equal(h5Style?.fontWeight, '300');
```

코드는 `font-weight` 태그의 `h6` 속성을 200으로 설정해야 합니다.

```js
const h6Element =document.querySelector('h6')
const h6Style = window.getComputedStyle(h6Element);
assert.equal(h6Style?.fontWeight, '200');
```

# --seed--

## --seed-contents--

```html
<style>
  h1 {
    font-size: 68px;

  }
  h2 {
    font-size: 52px;

  }
  h3 {
    font-size: 40px;

  }
  h4 {
    font-size: 32px;

  }
  h5 {
    font-size: 21px;

  }
  h6 {
    font-size: 14px;

  }
</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```

# --solutions--

```html
<style>
  h1 {
    font-size: 68px;
    font-weight: 800;
  }
  h2 {
    font-size: 52px;
    font-weight: 600;
  }
  h3 {
    font-size: 40px;
    font-weight: 500;
  }
  h4 {
    font-size: 32px;
    font-weight: 400;
  }
  h5 {
    font-size: 21px;
    font-weight: 300;
  }
  h6 {
    font-size: 14px;
    font-weight: 200;
  }
</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```
