---
id: 587d78a4367417b2b2512ad2
title: 제3색에 대해 배우기
challengeType: 0
forumTopicId: 301057
dashedName: learn-about-tertiary-colors
---

# --description--

컴퓨터 모니터와 기기 화면은 빨강, 초록, 파랑 빛의 양을 조합해 다양한 색을 만듭니다. 이것은 현대 색채 이론에서 RGB 가산 색 모델로 알려져 있습니다. 빨강(R), 초록(G), 파랑(B)은 기본색이라 부릅니다. 두 기본색을 섞으면 보조색인 시안(G + B), 마젠타(R + B), 노랑(R + G)이 만들어집니다. 이 색들은 보색 도전 과제에서 보셨습니다. 이 보조색들은 자신을 만든 데 사용되지 않은 기본색의 보색이며, 색상환에서 그 기본색과 반대 위치에 있습니다. 예를 들어, 마젠타는 빨강과 파랑으로 만들어지며 초록의 보색입니다.

제3색은 기본색과 그 보조색 이웃 중 하나를 조합해 만듭니다. 예를 들어, RGB 색 모델에서 빨강(기본색)과 노랑(보조색)을 섞으면 주황(제3색)이 됩니다. 이렇게 해서 단순한 색상환에 여섯 가지 색이 더해져 총 열두 가지 색이 됩니다.

디자인에서 조화로운 조합을 만드는 다양한 색 선택 방법이 있습니다. 제3색을 사용할 수 있는 예 중 하나가 분리 보색 색상 조합입니다. 이 조합은 기본색을 시작으로 그 보색의 양옆에 있는 두 색과 짝을 이룹니다. 세 가지 색은 디자인에서 강한 시각적 대비를 제공하지만 두 보색을 사용하는 것보다 더 은은합니다.

다음은 분리 보색 조합으로 만든 세 가지 색입니다:

<table><thead><tr><th>색상</th><th>16진수 코드</th></tr></thead><thead></thead><tbody><tr><td>주황</td><td>#FF7F00</td></tr><tr><td>시안</td><td>#00FFFF</td></tr><tr><td>라즈베리</td><td>#FF007F</td></tr></tbody></table>

# --instructions--

`background-color`, `orange`, `cyan` 클래스의 `raspberry` 속성을 각각 해당 색상으로 변경하세요. 색상 이름이 아닌 16진수 코드를 사용해야 합니다.

# --hints--

`div` 클래스를 가진 `orange` 요소는 주황색 `background-color`을 가져야 합니다.

```js
const orangeElement = document.querySelector('.orange');
const orangeStyle = window.getComputedStyle(orangeElement);
assert.equal(orangeStyle?.backgroundColor, 'rgb(255, 127, 0)');
```

`div` 클래스를 가진 `cyan` 요소는 시안색 `background-color`을 가져야 합니다.

```js
const cyanElement = document.querySelector('.cyan');
const cyanStyle = window.getComputedStyle(cyanElement);
assert.equal(cyanStyle?.backgroundColor, 'rgb(0, 255, 255)');
```

`div` 클래스를 가진 `raspberry` 요소는 라즈베리색 `background-color`을 가져야 합니다.

```js
const raspberryElement = document.querySelector('.raspberry');
const raspberryStyle = window.getComputedStyle(raspberryElement);
assert.equal(raspberryStyle?.backgroundColor, 'rgb(255, 0, 127)');
```

색상 클래스의 모든 `background-color` 값은 색상 이름이 아닌 16진수 코드여야 합니다.

```js
assert.notMatch(code, /background-color:\s(orange|cyan|raspberry)/);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .orange {
    background-color: #000000;
  }

  .cyan {
    background-color: #000000;
  }

  .raspberry {
    background-color: #000000;
  }

  div {
    height: 100px;
    width: 100px;
    margin-bottom: 5px;
  }
</style>

<div class="orange"></div>
<div class="cyan"></div>
<div class="raspberry"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .orange {
    background-color: #FF7F00;
  }

  .cyan {
    background-color: #00FFFF;
  }

  .raspberry {
    background-color: #FF007F;
  }

  div {
    height: 100px;
    width: 100px;
    margin-bottom: 5px;
  }
</style>
<div class="orange"></div>
<div class="cyan"></div>
<div class="raspberry"></div>
```
