---
id: 587d78a5367417b2b2512ada
title: 호버 시 요소 크기를 조절하는 CSS Transform scale 속성 사용하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLPJuM'
forumTopicId: 301077
dashedName: use-the-css-transform-scale-property-to-scale-an-element-on-hover
---

# --description--

`transform` 속성은 요소를 크기 조절, 이동, 회전, 기울이기 등 다양한 기능을 제공합니다. `:hover` 같은 가상 클래스와 함께 사용하면, `transform` 속성으로 요소에 쉽게 인터랙티브 효과를 추가할 수 있습니다.

사용자가 요소 위에 마우스를 올렸을 때 단락 요소를 원래 크기의 2.1배로 키우는 예시는 다음과 같습니다:

```css
p:hover {
  transform: scale(2.1);
}
```

**참고:** `div` 요소에 변형을 적용하면 해당 div 안에 포함된 자식 요소에도 영향을 줍니다.

# --instructions--

`hover`의 `div` 상태에 대한 CSS 규칙을 추가하고, 사용자가 요소 위에 마우스를 올렸을 때 `transform` 요소를 원래 크기의 1.1배로 키우도록 `div` 속성을 사용하세요.

# --hints--

사용자가 `div` 요소 위에 마우스를 올렸을 때 크기가 1.1배로 조절되어야 합니다.

```js
assert.match(code, /div:hover\s*?{\s*?transform:\s*?scale\(1\.1\);/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
    background: linear-gradient(
      53deg,
      #ccfffc,
      #ffcccf
    );
  }



</style>

<div></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
    background: linear-gradient(
      53deg,
      #ccfffc,
      #ffcccf
    );
  }
  div:hover {
    transform: scale(1.1);
  }
</style>
<div></div>
```
