---
id: 587d781e367417b2b2512aca
title: CSS 오프셋으로 상대 위치 지정 요소 이동하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bQEA4'
forumTopicId: 301065
dashedName: move-a-relatively-positioned-element-with-css-offsets
---

# --description--

`top` 또는 `bottom`와 `left` 또는 `right`의 CSS 오프셋은 문서의 일반 흐름에서 요소가 위치할 자리에서 얼마나 떨어져 오프셋할지 브라우저에 알려줍니다. 요소를 특정 위치에서 떨어뜨려 이동시키는 것으로, 참조된 쪽에서 반대 방향으로 요소가 이동합니다. 이전 도전 과제에서 보았듯이 `top` 오프셋을 사용하면 `h2`가 아래로 이동했습니다. 마찬가지로 `left` 오프셋을 사용하면 요소가 오른쪽으로 이동합니다.

# --instructions--

CSS 오프셋을 사용해 `h2`를 오른쪽으로 15픽셀, 위로 10픽셀 이동하세요.

# --hints--

코드는 `h2`를 위로 10px 상대 위치 지정하기 위해 CSS 오프셋을 사용해야 합니다. 즉, 요소가 원래 위치의 `bottom`에서 10px 떨어지도록 이동시키세요.

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.bottom, '10px');
```

코드는 `h2`를 오른쪽으로 15px 상대 위치 지정하기 위해 CSS 오프셋을 사용해야 합니다. 즉, 요소가 원래 위치의 `left`에서 15px 떨어지도록 이동시키세요.

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.left , '15px');
```

# --seed--

## --seed-contents--

```html
<head>
<style>
  h2 {
    position: relative;


  }
</style>
</head>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```

# --solutions--

```html
<head>
<style>
  h2 {
    position: relative;
    left: 15px;
    bottom: 10px;
  }
</style>
</head>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```
