---
id: 587d78a3367417b2b2512ad1
title: 보색에 대해 배우기
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MD3Tr'
forumTopicId: 301056
dashedName: learn-about-complementary-colors
---

# --description--

색상 이론과 디자인에 미치는 영향은 깊이 있는 주제이며, 다음 도전 과제에서는 기본만 다룹니다. 웹사이트에서 색상은 콘텐츠에 주목을 끌거나 감정을 불러일으키거나 시각적 조화를 만들 수 있습니다. 다양한 색상 조합을 사용하면 웹사이트의 모습이 크게 달라질 수 있으며, 콘텐츠와 어울리는 색상 팔레트를 고르는 데 많은 고민이 필요합니다.

색상환은 색상들이 서로 어떻게 관련되는지 시각화하는 데 유용한 도구입니다. 비슷한 색조는 이웃해 있고, 다른 색조는 더 멀리 떨어져 있는 원형입니다. 색상환에서 서로 마주 보는 두 색상을 보색이라고 합니다. 이들은 결합하면 서로를 "상쇄"하여 회색을 만들지만, 나란히 놓으면 더 생생해 보이고 강한 시각적 대비를 만듭니다.

보색과 그 헥스 코드 예시는 다음과 같습니다:

<blockquote>빨강 (#FF0000)과 시안 (#00FFFF)<br>초록 (#00FF00)과 마젠타 (#FF00FF)<br>파랑 (#0000FF)과 노랑 (#FFFF00)</blockquote>

이는 학교에서 배운 구식 RYB 색상 모델과 다릅니다. RYB 모델은 기본색과 보색이 다릅니다. 현대 색상 이론은 가산 RGB 모델(컴퓨터 화면처럼)과 감산 CMY(K) 모델(인쇄처럼)을 사용합니다.

온라인에는 색상의 보색을 찾는 옵션이 있는 다양한 색상 선택 도구가 있습니다.

**참고:** 색상은 페이지에 시각적 흥미를 더하는 강력한 방법입니다. 하지만 색상만으로 중요한 정보를 전달해서는 안 됩니다. 시각 장애가 있는 사용자는 그 내용을 이해하지 못할 수 있기 때문입니다. 이 문제는 적용된 접근성 도전 과제에서 더 자세히 다룹니다.

# --instructions--

`background-color`와 `blue` 클래스의 `yellow` 속성을 각각의 색상으로 변경하세요. 색상이 흰색 배경과 비교할 때와 서로 옆에 있을 때 어떻게 다르게 보이는지 주목하세요.

# --hints--

클래스 `div`를 가진 `blue` 요소는 파랑 색상의 `background-color`를 가져야 합니다.

```js
const blueElement = document.querySelector('.blue');
const blueStyle = window.getComputedStyle(blueElement); 
assert.equal(blueStyle?.backgroundColor, 'rgb(0, 0, 255)');
```

클래스 `div`를 가진 `yellow` 요소는 노랑 색상의 `background-color`를 가져야 합니다.

```js
const yellowElement = document.querySelector('.yellow');
const yellowStyle = window.getComputedStyle(yellowElement);
assert.equal(yellowStyle?.backgroundColor, 'rgb(255, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }
  .blue {
    background-color: #000000;
  }
  .yellow {
    background-color: #000000;
  }
  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>
<div class="blue"></div>
<div class="yellow"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }
  .blue {
    background-color: blue;
  }
  .yellow {
    background-color: yellow;
  }
  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>
<div class="blue"></div>
<div class="yellow"></div>
```
