---
id: 587d781e367417b2b2512acb
title: 절대 위치 지정으로 요소를 부모에 고정하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLJ7c3'
forumTopicId: 301060
dashedName: lock-an-element-to-its-parent-with-absolute-positioning
---

# --description--

CSS `position` 속성의 다음 옵션은 `absolute`으로, 요소를 부모 컨테이너에 상대적으로 고정합니다. `relative` 위치와 달리, 이 방법은 요소를 문서의 일반 흐름에서 제거하여 주변 항목들이 이를 무시하게 만듭니다. 위치를 조정하려면 CSS 오프셋 속성(위쪽 또는 아래쪽, 왼쪽 또는 오른쪽)을 사용하세요.

절대 위치 지정의 한 가지 미묘한 점은 가장 가까운 *위치가 지정된* 조상 요소에 상대적으로 고정된다는 것입니다. 부모 요소에 위치 규칙을 추가하는 것을 잊으면(보통 `position: relative;`를 사용해 설정합니다), 브라우저는 계속 상위 요소를 찾아가며 결국 `body` 태그를 기본값으로 사용합니다.

# --instructions--

`#searchbar` 요소를 `section` 부모의 오른쪽 상단에 고정하려면 `position`을 `absolute`로 선언하세요. 그리고 각각 50픽셀씩 `top`과 `right` 오프셋을 지정하세요.

# --hints--

`#searchbar` 요소는 `position`가 `absolute`로 설정되어야 합니다.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.position, 'absolute');
```

코드에서 `top` 요소에 `#searchbar` CSS 오프셋 50픽셀을 사용해야 합니다.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.top, '50px');
```

코드에서 `right` 요소에 `#searchbar` CSS 오프셋 50픽셀을 사용해야 합니다.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.right, '50px');
```

# --seed--

## --seed-contents--

```html
<style>
  #searchbar {



  }
  section {
    position: relative;
  }
</style>
<body>
  <h1>Welcome!</h1>
  <section>
    <form id="searchbar">
      <label for="search">Search:</label>
      <input type="search" id="search" name="search">
      <input type="submit" name="submit" value="Go!">
    </form>
  </section>
</body>
```

# --solutions--

```html
<style>
  #searchbar {
    position: absolute;
    top: 50px;
    right: 50px;
  }
  section {
    position: relative;
  }
</style>
<body>
  <h1>Welcome!</h1>
  <section>
    <form id="searchbar">
      <label for="search">Search:</label>
      <input type="search" id="search" name="search">
      <input type="submit" name="submit" value="Go!">
    </form>
  </section>
</body>
```
