---
id: 587d78a5367417b2b2512ad6
title: 점진적인 CSS 선형 그라디언트 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4dpt9'
forumTopicId: 301047
dashedName: create-a-gradual-css-linear-gradient
---

# --description--

HTML 요소에 색상을 적용하는 것은 단일 평면 색상에만 국한되지 않습니다. CSS는 요소에 색상 전환, 즉 그라디언트를 사용할 수 있는 기능을 제공합니다. 이는 `background` 속성의 `linear-gradient()` 함수를 통해 접근합니다. 일반적인 구문은 다음과 같습니다:

```css
background: linear-gradient(gradient_direction, color 1, color 2, color 3, ...);
```

첫 번째 인수는 색상 전환이 시작되는 방향을 지정합니다. 이는 각도로 표현할 수 있으며, `90deg`는 수평 그라디언트(왼쪽에서 오른쪽으로)를 만들고 `45deg`는 대각선 그라디언트(왼쪽 아래에서 오른쪽 위로)를 만듭니다. 다음 인수들은 그라디언트에 사용되는 색상의 순서를 지정합니다.

예시:

```css
background: linear-gradient(90deg, red, yellow, rgb(204, 204, 255));
```

# --instructions--

`linear-gradient()` 요소의 `div`에 `background`를 사용하고, 방향을 35도로 설정하여 색상을 `#CCFFFF`에서 `#FFCCCC`로 변경하세요.

# --hints--

`div` 요소는 지정된 방향과 색상을 가진 `linear-gradient` `background`을 가져야 합니다.

```js
const divElement = document.querySelector('div');
const divStyle = window.getComputedStyle(divElement); 
assert.match(divStyle?.background, /linear-gradient\(35deg, rgb\(204, 255, 255\), rgb\(255, 204, 204\)\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin: 50px auto;

  }

</style>

<div></div>
```

# --solutions--

```html
<style>
  div {
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin: 50px auto;
    background: linear-gradient(35deg, #CCFFFF, #FFCCCC);
  }
</style>
<div></div>
```
