---
id: 587d781e367417b2b2512ac9
title: 요소의 상대 위치 변경
challengeType: 0
videoUrl: 'https://scrimba.com/c/czVmMtZ'
forumTopicId: 301044
dashedName: change-an-elements-relative-position
---

# --description--

CSS는 각 HTML 요소를 자체 박스로 처리하며, 이를 보통 <dfn>CSS 박스 모델</dfn>이라고 합니다. 블록 레벨 항목은 자동으로 새 줄에서 시작합니다(제목, 문단, div 등 생각하세요). 반면 인라인 항목은 주변 콘텐츠 내에 위치합니다(이미지나 span처럼). 이렇게 요소의 기본 배치를 문서의 <dfn>정상 흐름</dfn>이라고 하며, CSS는 position 속성으로 이를 재정의할 수 있습니다.

요소의 position이 `relative`로 설정되면, 페이지 정상 흐름에서 현재 위치를 기준으로 CSS가 요소를 *상대적으로* 어떻게 이동할지 지정할 수 있습니다. 이는 `left` 또는 `right`, 그리고 `top` 또는 `bottom`의 CSS 오프셋 속성과 함께 사용됩니다. 이 속성들은 요소를 원래 위치에서 픽셀, 백분율, em 단위로 *얼마나 떨어뜨려* 이동할지 나타냅니다. 다음 예시는 문단을 아래쪽에서 10픽셀 떨어뜨려 이동합니다:

```css
p {
  position: relative;
  bottom: 10px;
}
```

요소의 position을 relative로 변경해도 정상 흐름에서 제거되지 않으므로, 주변 다른 요소들은 해당 항목이 기본 위치에 있는 것처럼 동작합니다.

**참고:** 위치 지정은 페이지 시각적 배치에 대해 많은 유연성과 강력한 제어를 제공합니다. 하지만 요소의 위치와 상관없이, 기본 HTML 마크업은 위에서 아래로 읽을 때 논리적으로 정리되어 있어야 한다는 점을 기억하세요. 이는 화면 읽기 프로그램 같은 보조 기기에 의존하는 시각 장애 사용자가 콘텐츠에 접근하는 방법입니다.

# --instructions--

`position`의 `h2`를 `relative`로 변경하고, CSS 오프셋을 사용해 정상 흐름에서 위치한 `top`에서 15픽셀 떨어지도록 이동하세요. 주변의 h1과 p 요소 위치에는 영향이 없음을 확인하세요.

# --hints--

`h2` 요소에 `position` 속성을 `relative`로 설정해야 합니다.

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.position, 'relative');
```

코드는 CSS 오프셋을 사용해 `h2`를 정상 위치에서 `top` 방향으로 15px 떨어지게 상대 위치 지정해야 합니다.

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.top, '15px');
```

# --seed--

## --seed-contents--

```html
<style>
  h2 {


  }
</style>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```

# --solutions--

```html
<style>
  h2 {
    position: relative;
    top: 15px;
  }
</style>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```
