---
id: 587d78a6367417b2b2512adc
title: CSS Transform 속성 skewY를 사용해 요소를 Y축을 따라 기울이기
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MZ2uB'
forumTopicId: 301075
dashedName: use-the-css-transform-property-skewy-to-skew-an-element-along-the-y-axis
---

# --description--

`skewX()` 함수가 선택한 요소를 주어진 각도로 X축을 따라 기울이는 것처럼, `skewY()` 속성도 요소를 Y(수직)축을 따라 기울입니다.

# --instructions--

`top` 속성을 사용해 id가 `transform`인 요소를 Y축을 따라 -10도 기울이세요.

# --hints--

id가 `top`인 요소는 Y축을 따라 -10도 기울여야 합니다.

```js
assert.match(code, /#top\s*?{\s*?.*?\s*?transform:\s*?skewY\(-10deg\);/g);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin: 50px auto;
  }
  #top {
    background-color: red;

  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>

<div id="top"></div>
<div id="bottom"></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin: 50px auto;
  }
  #top {
    background-color: red;
    transform: skewY(-10deg);
  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>
<div id="top"></div>
<div id="bottom"></div>
```
