---
id: 587d78a7367417b2b2512ae0
title: 버튼의 호버 상태를 변경하기 위해 CSS 애니메이션 사용하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4vZAa'
forumTopicId: 301073
dashedName: use-css-animation-to-change-the-hover-state-of-a-button
---

# --description--

CSS `@keyframes`를 사용해 버튼의 호버 상태에서 색상을 변경할 수 있습니다.

다음은 호버 시 이미지의 너비를 변경하는 예시입니다:

```html
<style>
  img {
    width: 30px;
  }
  img:hover {
    animation-name: width;
    animation-duration: 500ms;
  }

  @keyframes width {
    100% {
      width: 40px;
    }
  }
</style>

<img src="https://cdn.freecodecamp.org/curriculum/applied-visual-design/google-logo.png" alt="Google's Logo" />
```

# --instructions--

`ms`는 밀리초를 나타내며, 1000ms는 1초와 같습니다.

CSS `@keyframes`를 사용해 `background-color` 요소의 `button`를 변경하여 사용자가 요소 위에 마우스를 올렸을 때 `#4791d0`가 되도록 하세요. `@keyframes` 규칙에는 `100%`에 대한 항목만 있어야 합니다.

# --hints--

@keyframes 규칙은 `animation-name` 배경색을 사용해야 합니다.

```js
assert.match(code, /@keyframes\s+?background-color\s*?{/g);
```

`@keyframes` 아래에는 `background-color`를 100%에서 `#4791d0`로 변경하는 규칙이 하나 있어야 합니다.

```js
assert.match(code, /100%\s*?{\s*?background-color:\s*?#4791d0;\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }

  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
  }


</style>

<button>Register</button>
```

# --solutions--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }

  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
  }

  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```
