---
id: 587d78a7367417b2b2512adf
title: CSS @keyframes와 animation 속성 작동 방식 배우기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cakprhv'
forumTopicId: 301059
dashedName: learn-how-the-css-keyframes-and-animation-properties-work
---

# --description--

요소에 애니메이션을 적용하려면 animation 속성과 `@keyframes` 규칙에 대해 알아야 합니다. animation 속성은 애니메이션이 어떻게 동작할지 제어하고 `@keyframes` 규칙은 애니메이션 동안 어떤 일이 일어나는지 제어합니다. 총 8개의 animation 속성이 있습니다. 이 도전 과제에서는 가장 중요한 두 가지부터 간단히 다룹니다:

`animation-name`는 애니메이션의 이름을 설정하며, 나중에 `@keyframes`가 CSS에 어떤 규칙이 어떤 애니메이션과 연결되는지 알려줍니다.

`animation-duration`는 애니메이션의 지속 시간을 설정합니다.

`@keyframes`는 애니메이션 기간 동안 정확히 어떤 일이 일어나는지 지정하는 방법입니다. 이는 애니메이션 중 특정 "프레임"에 대해 0%에서 100%까지의 백분율로 CSS 속성을 지정하여 이루어집니다. 이를 영화에 비유하면, 0%의 CSS 속성은 시작 장면에서 요소가 어떻게 보이는지, 100%의 CSS 속성은 엔딩 크레딧 직전 요소가 어떻게 보이는지를 나타냅니다. 그런 다음 CSS가 주어진 기간 동안 요소를 전환하여 장면을 연출하는 마법을 적용합니다. 다음은 `@keyframes`와 animation 속성 사용법을 설명하는 예시입니다:

```css
#anim {
  animation-name: colorful;
  animation-duration: 3s;
}

@keyframes colorful {
  0% {
    background-color: blue;
  }
  100% {
    background-color: yellow;
  }
}
```

`anim` id를 가진 요소에 대해 위 코드 스니펫은 `animation-name`를 `colorful`로 설정하고 `animation-duration`을 3초로 설정합니다. 그런 다음 `@keyframes` 규칙이 `colorful`라는 이름의 animation 속성과 연결됩니다. 애니메이션 시작(0%) 시 색상을 파란색으로 설정하고 애니메이션 끝(100%)에는 노란색으로 전환됩니다. 시작과 끝 전환에만 제한되지 않고 0%에서 100% 사이의 어떤 백분율에도 요소의 속성을 설정할 수 있습니다.

# --instructions--

id가 `rect`인 요소에 대해 `animation-name`를 `rainbow`로, `animation-duration`를 4초로 설정하여 애니메이션을 만드세요. 다음으로 `@keyframes` 규칙을 선언하고 애니메이션 시작(`background-color`) 시 `0%`를 파란색, 애니메이션 중간(`50%`) 시 초록색, 애니메이션 끝(`100%`) 시 노란색으로 설정하세요.

# --hints--

id가 `rect`인 요소는 `animation-name` 속성을 `rainbow` 값으로 가져야 합니다.

```js
const rectElement = document.querySelector('#rect');
const rectStyle = window.getComputedStyle(rectElement);
assert.equal(rectStyle?.animationName, 'rainbow');
```

id가 `rect`인 요소는 `animation-duration` 속성을 4s 값으로 가져야 합니다.

```js
const rectElement = document.querySelector('#rect');
const rectStyle = window.getComputedStyle(rectElement);
assert.equal(rectStyle?.animationDuration, '4s');
```

`@keyframes` 규칙은 `animation-name`가 `rainbow`여야 합니다.

```js
assert.match(code ,/@keyframes\s+?rainbow\s*?{/g);
```

`@keyframes` 규칙은 `rainbow`에 대해 0%에서 `background-color`가 `blue`여야 합니다.

```js
assert.match(code,/0%\s*?{\s*?background-color:\s*?blue;\s*?}/gi);
```

`@keyframes` 규칙은 `rainbow`에 대해 50%에서 `background-color`가 `green`여야 합니다.

```js
assert.match(code,/50%\s*?{\s*?background-color:\s*?green;\s*?}/gi);
```

`@keyframes` 규칙은 rainbow에 대해 100%에서 `background-color`가 `yellow`여야 합니다.

```js
assert.match(code,/100%\s*?{\s*?background-color:\s*?yellow;\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
  }

  #rect {


  }




</style>
<div id="rect"></div>
```

# --solutions--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
    }
    50% {
      background-color: green;
    }
    100% {
      background-color: yellow;
    }
  }
</style>
<div id="rect"></div>
```
