---
id: 587d781b367417b2b2512abc
title: 텍스트의 background-color 속성 조정
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDqwA6'
forumTopicId: 301032
dashedName: adjust-the-background-color-property-of-text
---

# --description--

전체 배경이나 텍스트 색상을 조정해 전경을 쉽게 읽을 수 있게 하는 대신, 강조하려는 텍스트를 담고 있는 요소에 `background-color`를 추가할 수 있습니다. 이 도전 과제에서는 `rgba()`를 `hex` 코드나 일반 `rgb()` 대신 사용합니다.

<blockquote>rgba는 다음을 의미합니다:<br>  r = 빨강<br>  g = 초록<br>  b = 파랑<br>  a = 알파/불투명도 수준</blockquote>

RGB 값은 0부터 255까지 범위가 가능합니다. 알파 값은 완전히 불투명하거나 단색인 1부터 완전히 투명하거나 지운 상태인 0까지 범위가 가능합니다. 이 경우 불투명도를 조절할 수 있으므로 배경을 완전히 가릴 필요가 없어 `rgba()`를 사용하기 좋습니다.

이번 도전 과제에서는 `background-color: rgba(45, 45, 45, 0.1)`를 사용합니다. 불투명도 값이 0.1로 낮아 거의 투명한 어두운 회색을 만듭니다.

# --instructions--

텍스트를 더 돋보이게 하려면 `background-color` 요소의 `h4`를 주어진 `rgba()` 값으로 조정하세요.

또한 `h4`에 대해 `height` 속성을 제거하고 `padding`를 10px로 추가하세요.

# --hints--

코드는 `background-color` 요소에 `h4` 속성을 추가하고 `rgba(45, 45, 45, 0.1)`로 설정해야 합니다.

```js
assert(
  /(background-color|background):rgba\(45,45,45,0?\.1\)(;?}|;)/gi.test(
    code.replace(/\s/g, '')
  )
);
```

코드는 `padding` 요소에 `h4` 속성을 추가하고 10픽셀로 설정해야 합니다.

```js
const h4Element = document.querySelector("h4"); 
const h4style = window.getComputedStyle(h4Element);
assert.equal(h4style?.paddingTop, "10px");
assert.equal(h4style?.paddingRight, "10px");
assert.equal(h4style?.paddingBottom, "10px");
assert.equal(h4style?.paddingLeft, "10px");
```

`height` 요소의 `h4` 속성은 제거해야 합니다.

```js
const h4Element = document.querySelector("h4"); 
const h4style = window.getComputedStyle(h4Element);
assert.notEqual(h4style?.height, '25px');
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;


  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Alphabet</h4>
      <hr>
      <p><em>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
    padding: 10px;
    background-color: rgba(45, 45, 45, 0.1);
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Alphabet</h4>
      <hr>
      <p><em>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
