---
id: 587d78a6367417b2b2512ade
title: CSS와 HTML을 사용해 더 복잡한 모양 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPpz4fr'
forumTopicId: 301050
dashedName: create-a-more-complex-shape-using-css-and-html
---

# --description--

세계에서 가장 인기 있는 모양 중 하나는 하트 모양이며, 이 도전 과제에서는 순수한 CSS를 사용해 하트를 만듭니다. 하지만 먼저 `::before`와 `::after` 가상 요소를 이해해야 합니다. `::before`는 선택한 요소의 첫 번째 자식인 가상 요소를 만듭니다; `::after`는 선택한 요소의 마지막 자식인 가상 요소를 만듭니다. 다음 예제에서는 `::before` 가상 요소를 사용해 `heart` 클래스를 가진 요소에 사각형을 추가했습니다:

```css
.heart::before {
  content: "";
  background-color: yellow;
  border-radius: 25%;
  position: absolute;
  height: 50px;
  width: 70px;
  top: -50px;
  left: 5px;
}
```

`::before`와 `::after` 가상 요소가 제대로 작동하려면 `content` 속성이 정의되어 있어야 합니다. 이 속성은 보통 선택한 요소에 사진이나 텍스트 같은 것을 추가할 때 사용합니다. `::before`와 `::after` 가상 요소를 사용해 모양을 만들 때도 `content` 속성이 필요하지만, 빈 문자열로 설정합니다. 위 예제에서 `heart` 클래스를 가진 요소는 `::before` 가상 요소를 가지고 있으며, 높이와 너비가 각각 `50px`와 `70px`인 노란색 사각형을 만듭니다. 이 사각형은 25% `border-radius` 덕분에 둥근 모서리를 가지며, 요소의 왼쪽에서 `5px`, 위쪽에서 `50px`만큼 절대 위치로 배치됩니다.

# --instructions--

화면에 있는 요소를 하트 모양으로 바꾸세요. `.heart::after` 선택자에서 `background-color`를 `pink`로, `border-radius`를 50%로 변경하세요.

다음으로, `heart` 클래스를 가진 요소(단지 `heart`)를 선택하고 `transform` 속성을 채우세요. `rotate()` 함수를 -45도로 사용하세요.

마지막으로 `.heart::before` 선택자에서 `content` 속성을 빈 문자열로 설정하세요.

# --hints--

`background-color` 속성은 `.heart::after` 선택자에서 `pink`이어야 합니다.

```js
const heartAfter = code.match(/\.heart::after\s*{[\s\S]+?[^\}]}/g)?.[0];
assert.match(heartAfter, /({|;)\s*background-color\s*:\s*pink\s*(;|})/g);
```

`border-radius`는 `.heart::after` 선택자에서 50%이어야 합니다.

```js
assert.lengthOf(code.match(/border-radius\s*?:\s*?50%/gi),2);
```

`transform` 속성은 `heart` 클래스에서 `rotate()` 함수를 -45도로 설정해 사용해야 합니다.

```js
assert.match(code,/transform\s*?:\s*?rotate\(\s*?-45deg\s*?\)/gi);
```

`content`는 `.heart::before` 선택자에서 빈 문자열이어야 합니다.

```js
assert.match(code,/\.heart::before\s*?{\s*?content\s*?:\s*?("|')\1\s*?;/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: ;
  }
  .heart::after {
    background-color: blue;
    content: "";
    border-radius: 25%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart::before {
    content: ;
    background-color: pink;
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }
</style>
<div class="heart"></div>
```

# --solutions--

```html
<style>
  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: rotate(-45deg);
  }
  .heart::after {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart::before {
    content: "";
    background-color: pink;
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }
</style>
<div class="heart"></div>
```
