---
id: 587d7791367417b2b2512ab3
title: text-align 속성을 사용해 시각적 균형 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/c/c3b4EAp'
forumTopicId: 301053
dashedName: create-visual-balance-using-the-text-align-property
---

# --description--

이 커리큘럼 섹션은 응용 시각 디자인에 집중합니다. 첫 번째 도전 과제 그룹은 주어진 카드 레이아웃을 바탕으로 여러 핵심 원칙을 보여줍니다.

텍스트는 웹 콘텐츠에서 큰 비중을 차지하는 경우가 많습니다. CSS는 `text-align` 속성으로 텍스트를 정렬하는 여러 옵션을 제공합니다.

`text-align: justify;` 은 각 줄의 너비가 같도록 텍스트 간격을 조절합니다.

`text-align: center;` 는 텍스트를 가운데 정렬합니다.

`text-align: right;` 은 텍스트를 오른쪽 정렬합니다.

그리고 `text-align: left;` (기본값)는 텍스트를 왼쪽 정렬합니다.

# --instructions--

"Google"이라고 적힌 `h4` 태그의 텍스트를 가운데 정렬하세요. 그리고 Google이 어떻게 설립되었는지에 대한 정보를 담은 문단 태그는 양쪽 정렬하세요.

# --hints--

`h4` 태그에 text-align 속성을 사용해 `center` 로 설정해야 합니다.

```js
const h4Element =document.querySelector('h4')
const h4Style = window.getComputedStyle(h4Element);
assert.equal(h4Style?.textAlign, 'center');
```

`p` 태그에 text-align 속성을 사용해 `justify` 로 설정해야 합니다.

```js
const pElement =document.querySelector('p')
const pStyle = window.getComputedStyle(pElement);
assert.equal(pStyle?.textAlign, 'justify');
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {

  }
  p {

  }
  .links {
    margin-right: 20px;

  }
  .fullCard {
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at Stanford University.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
  }
  p {
    text-align: justify;
  }
  .links {
    margin-right: 20px;

  }
  .fullCard {
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at Stanford University.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
