---
id: 587d78a5367417b2b2512ad8
title: 배경 이미지로 미묘한 패턴을 추가해 질감 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cQdwJC8'
forumTopicId: 301052
dashedName: create-texture-by-adding-a-subtle-pattern-as-a-background-image
---

# --description--

배경에 질감과 흥미를 더해 더 돋보이게 만드는 한 가지 방법은 미묘한 패턴을 추가하는 것입니다. 핵심은 균형을 맞추는 것으로, 배경이 너무 튀어 전경을 가리지 않도록 해야 합니다. `background` 속성은 선택한 질감이나 패턴 이미지에 링크하기 위해 `url()` 함수를 지원합니다. 링크 주소는 소괄호 안에 따옴표로 감싸서 작성합니다.

# --instructions--

`https://cdn-media-1.freecodecamp.org/imgr/MJAkxbh.png`의 URL을 사용해 `background` 선택자로 전체 페이지의 `body`를 설정하세요.

# --hints--

`body` 요소에 주어진 링크가 포함된 `background`로 `url()` 속성을 설정해야 합니다.

```js
assert.match(code,/background(-image)?:\s*?url\(\s*("|'|)https:\/\/cdn-media-1\.freecodecamp\.org\/imgr\/MJAkxbh\.png\2\s*\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {

  }
</style>
```

# --solutions--

```html
<style>
  body {
    background: url("https://cdn-media-1.freecodecamp.org/imgr/MJAkxbh.png");
  }
</style>
```
