---
id: 587d78a3367417b2b2512ace
title: float 속성으로 요소를 왼쪽이나 오른쪽으로 밀기
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MDqu2'
forumTopicId: 301066
dashedName: push-elements-left-or-right-with-the-float-property
---

# --description--

다음 위치 지정 도구는 실제로 `position`를 사용하지 않고 요소의 `float` 속성을 설정합니다. 떠 있는 요소는 문서의 일반 흐름에서 제거되어 포함하는 부모 요소의 `left` 또는 `right` 쪽으로 밀립니다. 보통 떠 있는 요소가 차지할 수평 공간을 지정하기 위해 `width` 속성과 함께 사용합니다.

# --instructions--

주어진 마크업은 `section`와 `aside` 요소가 나란히 있는 2열 레이아웃으로 잘 작동합니다. `#left` 항목에 `float`를 `left`로, `#right` 항목에 `float`를 `right`로 지정하세요.

# --hints--

id가 `left`인 요소는 `float` 값으로 `left`를 가져야 합니다.

```js
const leftElement = document.querySelector('#left');
const leftStyle = window.getComputedStyle(leftElement);
assert.equal(leftStyle?.float, 'left');
```

id가 `right`인 요소는 `float` 값으로 `right`를 가져야 합니다.

```js
const rightElement = document.querySelector('#right');
const rightStyle = window.getComputedStyle(rightElement);
assert.equal(rightStyle?.float, 'right');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
    #left {

      width: 50%;
    }
    #right {

      width: 40%;
    }
    aside, section {
      padding: 2px;
      background-color: #ccc;
    }
  </style>
</head>
<body>
  <header>
    <h1>Welcome!</h1>
  </header>
  <section id="left">
    <h2>Content</h2>
    <p>Good stuff</p>
  </section>
  <aside id="right">
    <h2>Sidebar</h2>
    <p>Links</p>
  </aside>
</body>
```

# --solutions--

```html
<head>
  <style>
    #left {
      float: left;
      width: 50%;
    }
    #right {
      float: right;
      width: 40%;
    }
    aside, section {
      padding: 2px;
      background-color: #ccc;
    }
  </style>
</head>
<body>
  <header>
    <h1>Welcome!</h1>
  </header>
  <section id="left">
    <h2>Content</h2>
    <p>Good stuff</p>
  </section>
  <aside id="right">
    <h2>Sidebar</h2>
    <p>Links</p>
  </aside>
</body>
```
