---
id: 587d78a4367417b2b2512ad5
title: 색상의 톤 조절하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDJvT7'
forumTopicId: 301038
dashedName: adjust-the-tone-of-a-color
---

# --description--

CSS의 `hsl()` 옵션은 색상의 톤을 쉽게 조절할 수 있게 해줍니다. 순수한 색상에 흰색을 섞으면 그 색상의 틴트가 만들어지고, 검은색을 더하면 쉐이드가 만들어집니다. 또는 회색을 더하거나 틴트와 쉐이드를 모두 더하면 톤이 만들어집니다. `hsl()`의 's'와 'l'은 각각 채도와 밝기를 의미합니다. 채도 백분율은 회색의 양을 바꾸고, 밝기 백분율은 색상에 포함된 흰색 또는 검은색의 양을 결정합니다. 기본 색상을 좋아하지만 다양한 변형이 필요할 때 유용합니다.

# --instructions--

모든 요소는 기본 `background-color`가 `transparent`입니다. 현재 `nav` 요소는 뒤에 있는 요소가 `cyan`를 `background-color`로 설정했기 때문에 `cyan` 배경으로 보입니다. `background-color` 요소에 `nav`를 추가하여 같은 `cyan` 색상을 사용하지만 `80%` 채도와 `25%` 밝기 값을 적용해 톤과 쉐이드를 변경하세요.

# --hints--

`nav` 요소는 `background-color` 속성을 사용해 조절된 시안 톤의 `hsl()`를 가져야 합니다.

```js
// Computed style of hsl(180, 80%, 25%) results in rgb(13,115,115)
assert.equal(
  new __helpers.CSSHelp(document).getStyle('nav').getPropVal('background-color', true), 
  'rgb(13,115,115)'
)
```

# --seed--

## --seed-contents--

```html
<style>
  header {
    background-color: hsl(180, 90%, 35%);
    color: #FFFFFF;
  }

  nav {

  }

  h1 {
    text-indent: 10px;
    padding-top: 10px;
  }

  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }

  nav li {
    display: inline;
    margin-right: 20px;
  }

  a {
    text-decoration: none;
    color: inherit;
  }
</style>

<header>
  <h1>Cooking with FCC!</h1>
  <nav>
    <ul>
      <li><a href="#">Home</a></li>
      <li><a href="#">Classes</a></li>
      <li><a href="#">Contact</a></li>
    </ul>
  </nav>
</header>
```

# --solutions--

```html
<style>
  header {
    background-color: hsl(180, 90%, 35%);
    color: #FFFFFF;
  }

  nav {
    background-color: hsl(180, 80%, 25%);
  }

  h1 {
    text-indent: 10px;
    padding-top: 10px;
  }

  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }

  nav li {
    display: inline;
    margin-right: 20px;
  }

  a {
    text-decoration: none;
    color: inherit;
  }
</style>
<header>
  <h1>Cooking with FCC!</h1>
  <nav>
    <ul>
      <li><a href="#">Home</a></li>
      <li><a href="#">Classes</a></li>
      <li><a href="#">Contact</a></li>
    </ul>
  </nav>
</header>
```
