---
id: 587d78a8367417b2b2512ae3
title: 무한 애니메이션 횟수로 요소를 계속 애니메이션하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJDVfq'
forumTopicId: 301041
dashedName: animate-elements-continually-using-an-infinite-animation-count
---

# --description--

이전 도전 과제에서는 일부 애니메이션 속성과 `@keyframes` 규칙을 사용하는 방법을 다뤘습니다. 또 다른 애니메이션 속성은 `animation-iteration-count`으로, 애니메이션을 몇 번 반복할지 제어할 수 있습니다. 예시는 다음과 같습니다:

```css
animation-iteration-count: 3;
```

이 경우 애니메이션은 3번 실행한 후 멈추지만, 이 값을 `infinite`로 설정하면 애니메이션이 계속 실행되도록 만들 수 있습니다.

# --instructions--

오른쪽에서 공이 계속 튀도록 하려면 `animation-iteration-count` 속성을 `infinite`로 변경하세요.

# --hints--

`animation-iteration-count` 속성 값은 `infinite`이어야 합니다.

```js
const ballElement = document.querySelector('#ball');
const ballStyle = window.getComputedStyle(ballElement); 
assert.equal(ballStyle?.animationIterationCount, 'infinite');
```

# --seed--

## --seed-contents--

```html
<style>

  #ball {
    width: 100px;
    height: 100px;
    margin: 50px auto;
    position: relative;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: bounce;
    animation-duration: 1s;
    animation-iteration-count: 3;
  }

  @keyframes bounce{
    0% {
      top: 0px;
    }
    50% {
      top: 249px;
      width: 130px;
      height: 70px;
    }
    100% {
      top: 0px;
    }
  }
</style>
<div id="ball"></div>
```

# --solutions--

```html
<style>
  #ball {
    width: 100px;
    height: 100px;
    margin: 50px auto;
    position: relative;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: bounce;
    animation-duration: 1s;
    animation-iteration-count: infinite;
  }

  @keyframes bounce{
    0% {
      top: 0px;
    }
    50% {
      top: 249px;
      width: 130px;
      height: 70px;
    }
    100% {
      top: 0px;
    }
  }
</style>
<div id="ball"></div>
```
