---
id: 587d781d367417b2b2512ac8
title: 앵커 태그의 호버 상태 조정
challengeType: 0
videoUrl: 'https://scrimba.com/c/cakRGcm'
forumTopicId: 301035
dashedName: adjust-the-hover-state-of-an-anchor-tag
---

# --description--

이 도전 과제에서는 가상 클래스 사용법을 다룹니다. 가상 클래스는 선택자에 추가할 수 있는 키워드로, 요소의 특정 상태를 선택할 때 사용합니다.

예를 들어, 앵커 태그의 호버 상태에 대해 스타일을 변경하려면 `:hover` 가상 클래스 선택자를 사용합니다. 다음은 앵커 태그의 `color`를 호버 상태에서 빨간색으로 변경하는 CSS입니다:

```css
a:hover {
  color: red;
}
```

# --instructions--

코드 편집기에는 모든 `a` 태그를 검은색으로 스타일링하는 CSS 규칙이 있습니다. 사용자가 `a` 태그 위에 마우스를 올렸을 때 `color`가 파란색이 되도록 규칙을 추가하세요.

# --hints--

앵커 태그 `color`는 검은색으로 유지되어야 하며, `:hover` 상태에 대한 CSS 규칙만 추가하세요.

```js
const anchorElement = document.querySelector("a"); 
const anchorStyle = window.getComputedStyle(anchorElement);
assert.equal(anchorStyle?.color, 'rgb(0, 0, 0)');
```

앵커 태그는 호버 시 `color`가 파란색이어야 합니다.

```js
assert.match(code,
    /a:hover\s*?{\s*?color:\s*?(blue|rgba\(\s*?0\s*?,\s*?0\s*?,\s*?255\s*?,\s*?1\s*?\)|#00F|rgb\(\s*?0\s*?,\s*?0\s*?,\s*?255\s*?\))\s*?;\s*?}/gi
  );
```

# --seed--

## --seed-contents--

```html
<style>
  a {
    color: #000;
  }



</style>
<a href="https://freecatphotoapp.com/" target="_blank">CatPhotoApp</a>
```

# --solutions--

```html
<style>
  a {
    color: #000;
  }
  a:hover {
    color: rgba(0,0,255,1);
  }
</style>
<a href="https://freecatphotoapp.com/" target="_blank">CatPhotoApp</a>
```
