---
id: 587d78a6367417b2b2512add
title: CSS로 그래픽 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDWPs6'
forumTopicId: 301048
dashedName: create-a-graphic-using-css
---

# --description--

다양한 선택자와 속성을 조작해서 흥미로운 모양을 만들 수 있습니다. 시도하기 쉬운 모양 중 하나는 초승달 모양입니다. 이 도전 과제에서는 요소의 그림자를 설정하는 `box-shadow` 속성과 요소 모서리의 둥근 정도를 조절하는 `border-radius` 속성을 사용해야 합니다.

둥글고 투명한 객체를 만들고, 그림자는 약간 옆으로 오프셋된 선명한 그림자가 되도록 만드세요 — 그림자가 실제로 보이는 달 모양이 될 것입니다.

둥근 객체를 만들려면 `border-radius` 속성을 50% 값으로 설정해야 합니다.

이전 도전 과제에서 기억할 수 있듯이 `box-shadow` 속성은 `offset-x`, `offset-y`, `blur-radius`, `spread-radius` 그리고 `color` 값을 순서대로 받습니다. `blur-radius`와 `spread-radius` 값은 선택 사항입니다.

# --instructions--

편집기 안의 사각형 요소를 조작해서 달 모양을 만드세요. 먼저 `background-color`을 `transparent`로 변경한 다음, `border-radius` 속성을 50%로 설정해 원형 모양을 만드세요. 마지막으로 `box-shadow` 속성을 변경해 `offset-x`를 25px, `offset-y`를 10px, `blur-radius`를 0, `spread-radius`를 0, 그리고 `color`를 `blue`로 설정하세요.

# --hints--

`background-color` 속성의 값은 `transparent`로 설정해야 합니다.

```js
assert.match(code,/background-color:\s*?transparent;/gi);
```

`border-radius` 속성의 값은 `50%`로 설정해야 합니다.

```js
assert.match(code,/border-radius:\s*?50%;/gi);
```

`box-shadow` 속성의 값은 `offset-x`에 25px, `offset-y`에 10px, `blur-radius`에 0, `spread-radius`에 0, 마지막으로 `blue`에 `color`로 설정해야 합니다.

```js
assert.match(code,/box-shadow:\s*?25px\s+?10px\s+?0(px)?\s+?0(px)?\s+?blue\s*?;/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .center {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 0px;
    box-shadow: 25px 10px 10px 10px green;
  }

</style>
<div class="center"></div>
```

# --solutions--

```html
<style>
  .center {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    width: 100px;
    height: 100px;
    background-color: transparent;
    border-radius: 50%;
    box-shadow: 25px 10px 0 0 blue;
  }
</style>
<div class="center"></div>
```
