---
id: 587d78a9367417b2b2512ae9
title: 그래픽을 움직이기 위해 베지어 곡선 사용하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/c6bnRCK'
forumTopicId: 301071
dashedName: use-a-bezier-curve-to-move-a-graphic
---

# --description--

이전 도전 과제에서는 애니메이션이 처음에는 빨라졌다가 끝날 때 느려지는 변화를 설명하는 `ease-out` 키워드를 다뤘습니다. 오른쪽에서는 `ease-out` 키워드(파란 요소용)와 `linear` 키워드(빨간 요소용)의 차이를 보여줍니다. `ease-out` 키워드와 유사한 애니메이션 진행은 사용자 정의 3차 베지어 곡선 함수를 사용해 만들 수 있습니다.

일반적으로 `p1`와 `p2` 앵커 지점을 변경하면 다양한 베지어 곡선이 만들어지고, 이 곡선이 애니메이션이 시간에 따라 어떻게 진행되는지를 제어합니다. 다음은 `ease-out` 스타일을 모방하기 위해 값을 사용한 베지어 곡선 예시입니다:

```css
animation-timing-function: cubic-bezier(0, 0, 0.58, 1);
```

모든 `cubic-bezier` 함수는 `p0`에서 (0, 0)으로 시작해 `p3`에서 (1, 1)로 끝난다는 점을 기억하세요. 이 예시에서 곡선은 Y축을 따라 더 빠르게 움직입니다(0에서 시작해 `p1` y 값 0으로 갔다가 `p2` y 값 1로 이동) 반면 X축은 (0에서 시작해 `p1`에서 0, `p2`에서 0.58까지) 움직입니다. 결과적으로 애니메이션 요소의 변화가 해당 구간 애니메이션 시간보다 더 빠르게 진행됩니다. 곡선 끝부분에서는 x와 y 값 변화의 관계가 반전되어 y 값은 1에서 1로 유지(변화 없음)되고, x 값은 0.58에서 1로 이동해 애니메이션 변화가 애니메이션 지속 시간에 비해 느리게 진행됩니다.

# --instructions--

이 베지어 곡선 효과를 확인하려면, id가 `animation-timing-function`인 요소의 `red` 속성을 x1, y1, x2, y2가 각각 `cubic-bezier` 값으로 설정된 `0, 0, 0.58, 1` 함수로 변경하세요. 이렇게 하면 두 요소가 애니메이션을 비슷하게 진행합니다.

# --hints--

id가 `animation-timing-function`인 요소의 `red` 속성 값은 x1, y1, x2, y2 값이 각각 0, 0, 0.58, 1로 설정된 `cubic-bezier` 함수여야 합니다.

```js
const redElement = document.querySelector('#red');
const redStyle = window.getComputedStyle(redElement);
assert.equal(
  redStyle?.animationTimingFunction, 'cubic-bezier(0, 0, 0.58, 1)'
);
```

id가 `red`인 요소는 더 이상 `animation-timing-function`인 `linear` 속성을 가지면 안 됩니다.

```js
const redElement = document.querySelector('#red');
const redStyle = window.getComputedStyle(redElement);
assert.notEqual(redStyle?.animationTimingFunction, 'linear');
```

id가 `animation-timing-function`인 요소의 `blue` 속성 값은 변경하지 마세요.

```js
const blueElement = document.querySelector('#blue');
const blueStyle = window.getComputedStyle( blueElement);
const blueBallAnimation = __helpers.removeWhiteSpace(
  blueStyle?.animationTimingFunction
);
assert.isTrue(
  blueBallAnimation == 'ease-out' ||
    blueBallAnimation == 'cubic-bezier(0,0,0.58,1)'
);
```

# --seed--

## --seed-contents--

```html
<style>
  .balls{
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 27%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 56%;
    animation-timing-function: ease-out;
  }
  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id= "red"></div>
<div class="balls" id= "blue"></div>
```

# --solutions--

```html
<style>
  .balls{
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 27%;
    animation-timing-function: cubic-bezier(0, 0, 0.58, 1);
  }
  #blue {
    background: blue;
    left: 56%;
    animation-timing-function: ease-out;
  }
  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id= "red"></div>
<div class="balls" id= "blue"></div>
```
