---
id: 587d78a4367417b2b2512ad4
title: 색상의 색조 조절하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp38TZ'
forumTopicId: 301036
dashedName: adjust-the-hue-of-a-color
---

# --description--

색상은 색조, 채도, 밝기 등 여러 특성을 가집니다. CSS3는 이러한 특성을 직접 지정하여 색상을 선택하는 대안으로 `hsl()` 함수를 도입했습니다.

**색조**는 일반적으로 사람들이 '색'이라고 생각하는 것입니다. 왼쪽에 빨간색, 가운데에 초록색, 오른쪽에 파란색이 있는 색상 스펙트럼을 떠올리면, 색조는 이 선을 따라 색상이 위치하는 지점입니다. `hsl()`에서는 색상환 개념을 사용하여 원 위의 색상 각도를 0에서 360 사이 값으로 나타냅니다.

**채도**는 색상에 포함된 회색의 양입니다. 완전히 채도 높은 색상은 회색이 전혀 없고, 채도가 거의 없는 색상은 거의 완전한 회색입니다. 100%가 완전 채도를 의미하는 백분율로 표시합니다.

**밝기**는 색상에 포함된 흰색 또는 검은색의 양입니다. 0%(검은색)에서 100%(흰색)까지 백분율로 나타내며, 50%가 일반적인 색상입니다.

다음은 완전 채도와 일반 밝기를 가진 색상으로 `hsl()`를 사용하는 몇 가지 예입니다:

<table><thead><tr><th>색상</th><th>HSL</th></tr></thead><tbody><tr><td>red</td><td>hsl(0, 100%, 50%)</td></tr><tr><td>yellow</td><td>hsl(60, 100%, 50%)</td></tr><tr><td>green</td><td>hsl(120, 100%, 50%)</td></tr><tr><td>cyan</td><td>hsl(180, 100%, 50%)</td></tr><tr><td>blue</td><td>hsl(240, 100%, 50%)</td></tr><tr><td>magenta</td><td>hsl(300, 100%, 50%)</td></tr></tbody></table>

# --instructions--

`background-color`를 사용하여 클래스 이름(`div`, `green`, 또는 `cyan`)에 따라 각 `blue` 요소의 `hsl()`를 변경하세요. 세 가지 모두 완전 채도와 일반 밝기를 가져야 합니다.

# --hints--

코드는 녹색을 선언하기 위해 `hsl()` 함수를 사용해야 합니다.

```js
assert.match(code,/\.green\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

코드는 청록색을 선언하기 위해 `hsl()` 함수를 사용해야 합니다.

```js
assert.match(code,/\.cyan\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

코드는 파란색을 선언하기 위해 `hsl()` 함수를 사용해야 합니다.

```js
assert.match(code,/\.blue\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

클래스 `div`을 가진 `green` 요소는 녹색 `background-color`를 가져야 합니다.

```js
const greenElement = document.querySelector(".green");
const greenStyle = window.getComputedStyle(greenElement); 
assert.equal(greenStyle?.backgroundColor, 'rgb(0, 255, 0)');
```

클래스 `div`을 가진 `cyan` 요소는 청록색 `background-color`를 가져야 합니다.

```js
const cyanElement = document.querySelector(".cyan");
const cyanStyle = window.getComputedStyle(cyanElement); 
assert.equal(cyanStyle?.backgroundColor, 'rgb(0, 255, 255)');
```

클래스 `div`을 가진 `blue` 요소는 파란색 `background-color`를 가져야 합니다.

```js
const blueElement = document.querySelector(".blue");
const blueStyle = window.getComputedStyle(blueElement); 
assert.equal(blueStyle?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .green {
    background-color: #000000;
  }

  .cyan {
    background-color: #000000;
  }

  .blue {
    background-color: #000000;
  }

  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>

<div class="green"></div>
<div class="cyan"></div>
<div class="blue"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .green {
    background-color: hsl(120, 100%, 50%);
  }

  .cyan {
    background-color:   hsl(180, 100%, 50%);
  }

  .blue {
    background-color: hsl(240, 100%, 50%);
  }

  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>
<div class="green"></div>
<div class="cyan"></div>
<div class="blue"></div>
```
