---
id: 587d78a7367417b2b2512ae1
title: CSS 애니메이션을 사용해 움직임 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/c/c7amZfW'
forumTopicId: 301051
dashedName: create-movement-using-css-animation
---

# --description--

요소에 `position`가 지정되어 있을 때, 예를 들어 `fixed` 또는 `relative` 같은 경우, CSS 오프셋 속성인 `right`, `left`, `top`, 그리고 `bottom`를 애니메이션 규칙에서 사용해 움직임을 만들 수 있습니다.

아래 예시처럼, `top` 키프레임의 `50%` 속성을 50px로 설정해 아이템을 아래로 밀었다가, 첫 번째(`0%`)와 마지막(`100%`) 키프레임에서는 0px로 설정해 위로 움직이게 할 수 있습니다.

```css
@keyframes rainbow {
  0% {
    background-color: blue;
    top: 0px;
  }
  50% {
    background-color: green;
    top: 50px;
  }
  100% {
    background-color: yellow;
    top: 0px;
  }
}
```

# --instructions--

`div` 애니메이션에 수평 움직임을 추가하세요. `left` 오프셋 속성을 사용해 `@keyframes` 규칙에 다음을 추가하세요: 무지개가 `0%`에서 0픽셀, `50%`에서 25픽셀, `100%`에서 -25픽셀로 움직이도록 만드세요. 에디터에서 `top` 속성은 교체하지 마세요 — 애니메이션은 수직과 수평 움직임을 모두 가져야 합니다.

# --hints--

`@keyframes`에 대한 `0%` 규칙은 0px의 `left` 오프셋을 사용해야 합니다.

```js
assert.match(code,/[^50]0%\s*?{[\s\S]*?left:\s*?0px(;[\s\S]*?|\s*?)}/gi);
```

`@keyframes`에 대한 `50%` 규칙은 25px의 `left` 오프셋을 사용해야 합니다.

```js
assert.match(code,/50%\s*?{[\s\S]*?left:\s*?25px(;[\s\S]*?|\s*?)}/gi);
```

`@keyframes`에 대한 `100%` 규칙은 -25px의 `left` 오프셋을 사용해야 합니다.

```js
assert.match(code,/100%\s*?{[\s\S]*?left:\s*?-25px(;[\s\S]*?|\s*?)}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
    position: relative;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
      top: 0px;

    }
    50% {
      background-color: green;
      top: 50px;

    }
    100% {
      background-color: yellow;
      top: 0px;

    }
  }
</style>

<div id="rect"></div>
```

# --solutions--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
    position: relative;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
      top: 0px;
      left: 0px;
    }
    50% {
      background-color: green;
      top: 50px;
      left: 25px;
    }
    100% {
      background-color: yellow;
      top: 0px;
      left: -25px;
    }
  }
</style>
<div id="rect"></div>
```
