---
id: 58a7a6ebf9a6318348e2d5aa
title: 애니메이션의 채우기 모드 수정
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJDmcE'
forumTopicId: 301064
dashedName: modify-fill-mode-of-an-animation
---

# --description--

좋아요, 하지만 아직 제대로 작동하지 않습니다. `500ms`가 지나간 후 애니메이션이 리셋되어 버튼이 원래 색상으로 돌아가는 것을 확인하세요. 버튼이 강조된 상태로 유지되길 원합니다.

이것은 `animation-fill-mode` 속성을 `forwards`로 설정하면 됩니다. `animation-fill-mode`는 애니메이션이 끝난 후 요소에 적용되는 스타일을 지정합니다. 다음과 같이 설정할 수 있습니다:

```css
animation-fill-mode: forwards;
```

# --instructions--

사용자가 버튼 위에 마우스를 올렸을 때 버튼이 강조된 상태로 유지되도록 `animation-fill-mode`의 `button:hover` 속성을 `forwards`로 설정하세요.

# --hints--

`button:hover`는 값이 `animation-fill-mode`인 `forwards` 속성을 가져야 합니다.

```js
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-fill-mode\s*?:\s*?forwards\s*?;[\s\S]*}/gi);
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-name\s*?:\s*?background-color\s*?;[\s\S]*}/gi);
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-duration\s*?:\s*?500ms\s*?;[\s\S]*}/gi );
```

# --seed--

## --seed-contents--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
    /* Only change code below this line */

    /* Only change code above this line */
  }
  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```

# --solutions--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
    animation-fill-mode: forwards;
  }
  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```
