---
id: 69b2713ff6c21fed2d3ad88b
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Merge sort é um algoritmo de ordenação que usa o princípio de dividir e conquistar para ordenar coleções de dados. Ou seja, ele "divide" uma coleção em subpartes menores e "conquista" as subpartes ordenando-as independentemente, depois mescla as subpartes ordenadas.

Comece criando uma função chamada `mergeSort` que recebe o parâmetro `array`.

# --hints--

Você deve declarar uma função chamada `mergeSort` com um único parâmetro `array`.

```js
assert.isFunction(mergeSort);
const regex = __helpers.functionRegex('mergeSort', ['array']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
