---
id: 69b279bae0f20eda63163cfd
title: passo 19
challengeType: 1
dashedName: step-19
---

# --description--

Finalmente, você deve copiar quaisquer valores restantes da metade direita para `array`.

Para fazer isso, adicione um loop `while` que execute enquanto `rightArrayIndex` for menor que `rightPart.length`, atribua `rightPart[rightArrayIndex]` a `array[sortedIndex]` e incremente tanto `rightArrayIndex` quanto `sortedIndex`.

# --hints--

Você deve criar um loop `while` com a condição `rightArrayIndex < rightPart.length`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

O corpo do loop deve mover `rightPart[rightArrayIndex]` para `array[sortedIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?/
);
```

Você deve incrementar tanto `rightArrayIndex` quanto `sortedIndex` após a atribuição.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*rightArrayIndex\s*(\+\+|\+=\s*1)\s*;?[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
