---
id: 69b279bae0f20eda63163cfe
title: Passo 20
challengeType: 1
dashedName: step-20
---

# --description--

Antes de testar a função `mergeSort()`, você precisa criar um caso base que pare a execução da função quando o comprimento de `array` for menor ou igual a `1`.

Esse caso base vai interromper a chamada recursiva. Sem ele, a operação de ordenação por mesclagem continuaria a ser executada mesmo quando a lista já estivesse ordenada ou não tivesse elementos.

Logo após a declaração da função, crie uma declaração `if` que verifique se `array.length` é menor ou igual a 1. Dentro do corpo da declaração `if`, adicione uma declaração `return` para parar a execução da função.

# --hints--

Você deve criar uma declaração `if` com a condição `array.length <= 1`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{/
);
```

Você deve ter uma declaração `return` no corpo da declaração `if`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{[\s\S]*return\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }
  
  while (rightArrayIndex < rightPart.length) {
    array[sortedIndex] = rightPart[rightArrayIndex];
    rightArrayIndex += 1;
    sortedIndex += 1;
  }
}
```
