---
id: 69b279bae0f20eda63163cf8
title: passo 14
challengeType: 1
dashedName: step-14
---

# --description--

As declarações `if` e `else` que você criou nos passos anteriores vão atribuir elementos ao array ordenado.

Cada elemento atribuído ao array ordenado ocupa um índice na lista. Então você precisa passar para o próximo índice no array ordenado após cada atribuição.

Abaixo do bloco `if`/`else`, mas ainda dentro do loop `while`, incremente `sortedIndex` por `1`.

# --hints--

Você deve incrementar o valor atual de `sortedIndex` por `1` depois da cláusula `else`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
