---
id: 69b278d82f0e514f32acb724
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Depois de fatiar a metade esquerda, você ainda precisa do lado direito para completar o passo de mesclar.

Use a sintaxe de slice para extrair a metade direita de `array` e atribuí-la a uma variável chamada `rightPart`.

# --hints--

Você deve ter uma variável `rightPart` dentro da função `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightPart\s*=/
);
```

Você deve usar a sintaxe de slice para extrair a metade direita do array e atribuí-la à variável `rightPart`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /rightPart\s*=\s*array\.slice\(\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
