---
id: 69b279bae0f20eda63163cf3
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

Dentro do loop, comece a comparar os elementos atuais de cada metade para que você possa mesclar o menor valor primeiro.

Adicione um bloco `if` que verifique se o elemento atual em `leftPart` é menor que o elemento atual em `rightPart`. 
# --hints--

A condição `if` deve verificar se `leftPart[leftArrayIndex]` é menor que `rightPart[rightArrayIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*<\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
