---
id: 69b279bae0f20eda63163cf2
title: passo 8
challengeType: 1
dashedName: step-8
---

# --description--

Em seguida, você precisa criar um loop que continue enquanto houver elementos restantes em ambos `leftPart` e `rightPart`.

Para isso, crie um loop `while` com duas condições: uma que verifica se o `leftArrayIndex` é menor que o comprimento de `leftPart` e outra condição que verifica se `rightArrayIndex` é menor que o comprimento de `rightPart`.

# --hints--

Você deve criar um loop `while` que verifica se `leftArrayIndex` é menor que o comprimento de `leftPart` e se `rightArrayIndex` é menor que o comprimento de `rightPart`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
