---
id: 69b2762b17d23fbee40a6c00
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

O algoritmo de merge sort realiza três ações:

- Divida uma sequência não ordenada de itens em subpartes
- Ordene os itens nas subpartes
- Mesclar as subpartes ordenadas

O acima acontece recursivamente até que as subpartes sejam mescladas na sequência ordenada completa. Comece dividindo a sequência.

Primeiro, dentro da função `mergeSort`, declare uma variável `middlePoint` usando `const` e atribua a ela `Math.floor(array.length / 2)` para que você possa dividir a sequência ao meio antes de chamar a recursão.

# --hints--

Você deve declarar uma variável `middlePoint` e atribuir a ela `Math.floor(array.length / 2)`.

```js
const explorer = await __helpers.Explorer(code);

const { middlePoint } = explorer.allFunctions.mergeSort.variables;

assert.exists(
  middlePoint,
  'middlePoint should be declared inside mergeSort'
);

assert.isTrue(
  middlePoint.matches('const middlePoint = Math.floor(array.length / 2);'),
  'middlePoint should be declared with const and assigned Math.floor(array.length / 2)'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
