---
id: 69b279bae0f20eda63163cef
title: passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Agora que você dividiu a lista `array` em duas listas separadas, continue dividindo cada lista até que cada elemento esteja sozinho em sua própria lista. Uma lista com um único número está sempre ordenada.

Para fazer isso, chame recursivamente `mergeSort` dentro da sua função e passe `leftPart` como argumento para a chamada.

# --hints--

Você deve chamar `mergeSort(leftPart)` dentro da mesma função `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*leftPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
