---
id: 69b279bae0f20eda63163cf9
title: Passo 15
challengeType: 1
dashedName: step-15
---

# --description--

O loop `while` que você criou compara um elemento de `leftPart` com outro em `rightPart`, depois adiciona o elemento menor à lista principal `array`.

Ele continuará essa operação até não haver mais elementos para comparar. Mas `leftPart` ainda pode ter elementos sobrando enquanto `rightPart` não tem nenhum, e vice-versa.

Crie outro loop `while` que execute quando `leftArrayIndex` for menor que o comprimento de `leftPart`. Nos próximos passos, você vai usá-lo para copiar os elementos restantes em `leftPart` para a lista `array`.

# --hints--

Você deve criar um loop `while` com a condição `leftArrayIndex < leftPart.length`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
