---
id: 69b279bae0f20eda63163cf6
title: Passo 12
challengeType: 1
dashedName: step-12
---

# --description--

Em um passo anterior, você atribuiu o elemento em `leftPart` à lista `array` porque ele era menor. Mas isso nem sempre será o caso. Em alguns casos de comparação, o elemento à direita pode ser menor.

Crie uma cláusula `else` para executar quando o elemento em `leftPart` não for menor que o elemento em `rightPart`.

Dentro do bloco `else`, atribua `rightPart[rightArrayIndex]` a `array[sortedIndex]`.

# --hints--

Você deve criar um bloco `else` que contenha `array[sortedIndex] = rightPart[rightArrayIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /else\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
}
```
