---
id: 69b279bae0f20eda63163cf1
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Agora é hora de ordenar e mesclar as listas (`leftPart` e `rightPart`) na `array` original.

Você pode fazer isso comparando os elementos de ambas as listas e mesclando o elemento menor na lista principal. Você vai fazer essa comparação para todos os índices em `leftPart` e `rightPart`.

Crie três variáveis: `leftArrayIndex`, `rightArrayIndex` e `sortedIndex`, e defina seus valores como `0`. Essas variáveis vão ajudar você a acompanhar cada índice durante o processo de ordenação.

# --hints--

Você deve declarar uma variável `leftArrayIndex` inicializada com `0`.

```js
const explorer = await __helpers.Explorer(code);
const { leftArrayIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  leftArrayIndex,
  'You should declare a leftArrayIndex variable inside mergeSort'
);
assert.isTrue(
  leftArrayIndex.value.matches('0'),
  'leftArrayIndex should be initialized to 0'
);
```

Você deve declarar uma variável `rightArrayIndex` inicializada com `0`.

```js
const explorer = await __helpers.Explorer(code);
const { rightArrayIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  rightArrayIndex,
  'You should declare a rightArrayIndex variable inside mergeSort'
);
assert.isTrue(
  rightArrayIndex.value.matches('0'),
  'rightArrayIndex should be initialized to 0'
);
```

Você deve declarar uma variável `sortedIndex` inicializada com `0`.

```js
const explorer = await __helpers.Explorer(code);
const { sortedIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  sortedIndex,
  'You should declare a sortedIndex variable inside mergeSort'
);
assert.isTrue(
  sortedIndex.value.matches('0'),
  'sortedIndex should be initialized to 0'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
