---
id: 69b278519ef46b2fd8bf1546
title: passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Você já tem um valor `middlePoint`. Use-o para dividir o array em duas metades e deixe a parte esquerda pronta para ordenação.

Use a sintaxe slice para extrair a metade esquerda de `array` e atribuí-la a uma variável chamada `leftPart`.

# --hints--

Você deve declarar uma variável `leftPart` dentro da função `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftPart\s*=/
);
```

Você deve atribuir `array.slice(0, middlePoint)` a `leftPart`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /leftPart\s*=\s*array\.slice\(\s*0\s*,\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
