---
id: 69b279bae0f20eda63163cf4
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

Quando a condição `if` avalia para `true`, isso significa que o elemento na lista `leftPart` é menor que o elemento com o qual está sendo comparado na lista `rightPart`.

Nesse caso, você pode colocar o elemento encontrado em `leftArrayIndex` dentro de `leftPart` no array ordenado.

Dentro do bloco `if`, atribua `leftPart[leftArrayIndex]` a `array[sortedIndex]`.

# --hints--

Você deve atribuir `leftPart[leftArrayIndex]` a `array[sortedIndex]` dentro do `if`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /array\s*\[\s*sortedIndex\s*\]\s*=\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
