---
id: 69b279bae0f20eda63163cf0
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Na parte inferior do corpo da sua função, faça a chamada da função `mergeSort()` novamente. Desta vez, passe `rightPart` como argumento para a chamada da função.

# --hints--

Você deve chamar `mergeSort()` com o argumento `rightPart` na parte inferior do corpo da sua função.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*rightPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
