---
id: 69b278519ef46b2fd8bf1546
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Você já tem um valor `middlePoint`. Use-o para dividir o array em duas metades e deixe a parte esquerda pronta para ordenação.

Use a sintaxe slice para extrair a metade esquerda de `array` e atribuí-la a uma variável chamada `leftPart`.

# --hints--

Você deve declarar uma variável `leftPart` dentro da função `mergeSort`.

```js
const explorer = await __helpers.Explorer(code);
const { leftPart } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  leftPart,
  'You should declare a leftPart variable inside mergeSort'
);
```

Você deve atribuir `array.slice(0, middlePoint)` a `leftPart`.

```js
const explorer = await __helpers.Explorer(code);
const { leftPart } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  leftPart,
  'You should declare a leftPart variable inside mergeSort'
);
assert.isTrue(
  leftPart.value.matches('array.slice(0, middlePoint)'),
  'You should assign array.slice(0, middlePoint) to leftPart'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
