---
id: 69b2713ff6c21fed2d3ad88b
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Merge sort é um algoritmo de ordenação que usa o princípio de dividir e conquistar para ordenar coleções de dados. Ou seja, ele "divide" uma coleção em subpartes menores e "conquista" as subpartes ordenando-as independentemente, depois mescla as subpartes ordenadas.

Comece criando uma função chamada `mergeSort` que recebe o parâmetro `array`.

# --hints--

Você deve declarar uma função chamada `mergeSort` com um único parâmetro `array`.

```js
assert.isFunction(mergeSort);

const explorer = await __helpers.Explorer(code);
const mergeSortFn = explorer.allFunctions.mergeSort;
assert.exists(mergeSortFn, 'You should declare a mergeSort function');

const params = mergeSortFn.parameters;
assert.lengthOf(
  params,
  1,
  'mergeSort should take a single parameter'
);
assert.isTrue(
  params[0].matches('array'),
  'The parameter of mergeSort should be named array'
);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
