---
id: 66a207974c806a19d6607073
title: بناء تطبيق جرد الكتب
challengeType: 25
dashedName: build-a-book-inventory-app
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك عنصر `h1` نصّه `Book Inventory`.
1. يجب أن يكون لديك عنصر `table` بأعمدة بأسماء `Title` و`Author` و`Category` و`Status` و`Rate`.
1. يجب أن يحتوي جدولك على أربعة صفوف على الأقل، الأول لعناوين الأعمدة والبقية مملوءة بمعلومات الكتب.
1. يجب أن يحمل كل صف من صفوف الجدول داخل جسم الجدول الصنف `read` أو `to-read` أو `in-progress`.
1. يجب أن تحتوي عناصر `td` في العمود `Status` على عنصر `span` له `class` بقيمة `status` يحيط بالنص `Read` أو `To Read` أو `In Progress`، حسب صنف ذلك الصف.
1. يجب أن تحتوي عناصر `td` في العمود `Rate` على عنصر `span` له `class` بقيمة `rate` يغلّف ثلاثة عناصر `span` فارغة.
1. يجب أن تحمل عناصر `.rate` الموضوعة داخل صفوف `.read` صنفًا إضافيًا بقيمة `one` أو `two` أو `three`، حسب التقييم الشخصي. وينبغي أن تأتي هذه القيمة بعد `rate`.
1. يجب أن تنشئ ثلاثة محدِّدات خصائص لاستهداف العناصر ذات الصنف `read` و`to-read` و`in-progress`، واضبط خاصية `background-image` لها لاستخدام `linear-gradient` من اختيارك.
1. يجب أن تضبط خاصية `display` لكل عنصر `span` على `inline-block`.
1. يجب أن تستخدم محدِّد خصائص لاستهداف عناصر `span` ذات الصنف `status` التي تنحدر من عناصر `tr` ذات الصنف `to-read`، واضبط خاصيتي `border` و`background-image` لها.
1. يجب أن تستخدم محدِّد خصائص لاستهداف عناصر `span` ذات الصنف `status` التي تنحدر من عناصر `tr` ذات الصنف `read`، واضبط خاصيتي `border` و`background-image` لها.
1. يجب أن تستخدم محدِّد خصائص لاستهداف عناصر `span` ذات الصنف `status` التي تنحدر من عناصر `tr` ذات الصنف `in-progress`، واضبط خاصيتي `border` و`background-image` لها.
1. يجب أن تستخدم محدِّد خصائص لاستهداف عناصر `span` ذات الصنف `status` وعناصر `span` التي تبدأ قيمة صنفها بـ `rate`، واضبط خصائص `height` و`width` و`padding` لها.
1. يجب أن تستخدم محدِّد خصائص لاستهداف عناصر `span` التي هي أبناء مباشرون لعناصر `span` التي تبدأ قيمة `class` لها بـ `rate`، واضبط خصائص `border` و`border-radius` و`margin` و`height` و`width` و`background-color` لها.
1. يجب أن تستخدم محدِّد خصائص لاستهداف أول منحدر من عناصر `span` التي تتضمن الكلمة `one` كجزء من قيمة `class` لها، واضبط خاصية `background-image` لها لاستخدام `linear-gradient`.
1. يجب أن تستخدم محدِّد خصائص لاستهداف أول منحدرين من عناصر `span` التي تتضمن الكلمة `two` كجزء من قيمة `class` لها، واضبط خاصية `background-image` لها لاستخدام `linear-gradient`.
1. يجب أن تستخدم محدِّد خصائص لاستهداف عناصر `span` الثلاثة المنحدرة من عناصر `span` التي تتضمن الكلمة `three` كجزء من قيمة `class` لها، واضبط خاصية `background-image` لها لاستخدام `linear-gradient`.

# --hints--

يجب أن يكون لديك عنصر `h1` يحتوي على النص `Book Inventory`.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'Book Inventory');
```

يجب أن يكون لديك عنصر `h1` واحد فقط.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

يجب أن يكون لديك عنصر `table`.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

يجب أن يحتوي `table` على عنصر `thead` واحد وعنصر `tbody` واحد بداخله.

```js
assert.lengthOf(document.querySelectorAll("table > thead"), 1);
assert.lengthOf(document.querySelectorAll("table > tbody"), 1);
```

داخل `thead` يجب أن يكون هناك عنصر `tr` واحد يحتوي على 5 عناصر `th`.

```js
assert.lengthOf(document.querySelectorAll("table > thead > tr"), 1);
assert.lengthOf(document.querySelectorAll("table > thead > tr > th"), 5);
```

يجب أن يحتوي العمود الأول على النص `Title` كعنوان.

```js
assert.equal(document.querySelectorAll('th')[0]?.textContent.trim(), 'Title');
```

يجب أن يحتوي العمود الثاني على النص `Author` كعنوان.

```js
assert.equal(document.querySelectorAll('th')[1]?.textContent.trim(), 'Author');
```

يجب أن يحتوي العمود الثالث على النص `Category` كعنوان.

```js
assert.equal(document.querySelectorAll('th')[2]?.textContent.trim(), 'Category');
```

يجب أن يحتوي العمود الرابع على النص `Status` كعنوان.

```js
assert.equal(document.querySelectorAll('th')[3]?.textContent.trim(), 'Status');
```

يجب أن يحتوي العمود الخامس على النص `Rate` كعنوان.

```js
assert.equal(document.querySelectorAll('th')[4]?.textContent.trim(), 'Rate');
```

يجب أن يحتوي جدولك على أربعة صفوف على الأقل.

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
```

يجب أن يحتوي كل صف دائمًا على 5 أعمدة.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isNotEmpty(rows);
for (const row of rows) {
  assert.lengthOf(row.children, 5);
}
```

يجب أن يحتوي كل صف جدول باستثناء صف العنوان على أحد الأصناف `read`، `to-read`، أو `in-progress`.

```js
const rows = document.querySelectorAll('tr');
assert.isNotEmpty(rows);
for (let i = 1; i < rows.length; i++) {
  const classList = [...rows[i].classList];
  const currentProgress = classList[0]; 
  assert.oneOf(currentProgress, ['read', 'to-read', "in-progress"]); 
}
```

يجب أن تحتوي عناصر `td` في عمود `Status` على عنصر `span`.

```js
const statusColumnData = document.querySelectorAll('td:nth-child(4)');
assert.isNotEmpty(statusColumnData);
for (let e of statusColumnData) {
  assert.equal(e?.children[0]?.tagName, 'SPAN');
}
```

يجب أن يحمل كل عنصر `span` في عمود `Status` الصنف `status`.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
assert.isNotEmpty(statusSpans);
for (let e of statusSpans) {
    assert.isTrue(e.classList.contains('status'));
}
```

يجب أن يحتوي كل عنصر `.status` على النص `Read`، `To Read`، أو `In Progress`، حسب صنف صفه.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
const rows = Array.from(document.querySelectorAll('tr')).slice(1);
assert.isNotEmpty(statusSpans);
for (let i = 0; i < rows.length; i++) {
  switch (statusSpans[i]?.innerText.trim()) {
        case 'Read':
            assert.isTrue(rows[i].classList.contains('read'));
            break;
        case 'To Read':
            assert.isTrue(rows[i].classList.contains('to-read'));
            break;
        case 'In Progress':
            assert.isTrue(rows[i].classList.contains('in-progress'));
            break;
        default:
            assert.fail(); 
    }
}
```

يجب أن تحتوي عناصر `td` في عمود `Rate` على عنصر `span`.

```js
const rateColumnData = document.querySelectorAll('tr td:last-child');
assert.isNotEmpty(rateColumnData);
for (let e of rateColumnData) {
    assert.equal(e.children[0]?.tagName, 'SPAN')
}
```

يجب أن يحمل كل عنصر `span` هو ابن مباشر لعنصر `td` في عمود `Rate` الصنف `rate` كأول صنف.

```js
const rateSpans = document.querySelectorAll('tr td:last-child > span:first-child');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.classList[0], 'rate');
}
```

يجب أن يحتوي كل عنصر `.rate` على ثلاثة عناصر `span` فارغة.

```js
const rateSpans = document.getElementsByClassName('rate');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.children.length, 3);
    for (let child of e.children) {
        assert.equal(child.tagName, 'SPAN');
        assert.equal(child.innerText.length, 0);
    }
}
```

يجب أن تحتوي عناصر `.rate` الموجودة داخل صفوف `.read` على صنف إضافي بعد صنف `rate` بقيمة `one` أو `two` أو `three`، حسب التقييم الشخصي.

```js
const readBooksRates = document.querySelectorAll('.read .rate');
assert.isNotEmpty(readBooksRates);
for (let e of readBooksRates) {
    assert.oneOf(e.classList[1], ['one', 'two', 'three']);
}
```

يجب أن يكون لديك محدد خاصية لاستهداف الصفوف التي تحمل الصنف `read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"]'));
```

يجب استخدام محدد خاصية لاستهداف الصفوف التي تحمل الصنف `read` وتعيين خاصية `background-image` لها باستخدام تدرج خطي من اختيارك.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="read"]')?.backgroundImage, 'linear-gradient(');
```

يجب أن يكون لديك محدد خاصية لاستهداف الصفوف التي تحمل الصنف `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]'));
```

يجب استخدام محدد خاصية لاستهداف الصفوف التي تحمل الصنف `to-read` وتعيين خاصية `background-image` لها باستخدام تدرج خطي من اختيارك.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]')?.backgroundImage, 'linear-gradient(');
```

يجب أن يكون لديك محدد خاصية لاستهداف الصفوف التي تحمل الصنف `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]'));
```

يجب استخدام محدد خاصية لاستهداف الصفوف التي تحمل الصنف `in-progress` وتعيين خاصية `background-image` لها باستخدام تدرج خطي من اختيارك.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]')?.backgroundImage, 'linear-gradient(');
```

يجب تعيين خاصية `display` لكل عنصر `span` إلى `inline-block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('span')?.getPropVal('display'), 'inline-block');
```

يجب أن يكون لديك محدد خاصية لاستهداف عناصر `span` التي تحمل الصنف `status` والتي هي عناصر تابعة لـ `tr` التي تحمل الصنف `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]'));
```

يجب استخدام محدد خاصية لاستهداف عناصر `span` التي تحمل الصنف `status` والتي هي عناصر تابعة لـ `tr` التي تحمل الصنف `to-read` وتعيين خاصية `border` لها.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.border);
```

يجب استخدام محدد خاصية لاستهداف عناصر `span` التي تحمل الصنف `status` والتي هي عناصر تابعة لـ `tr` التي تحمل الصنف `to-read` وتعيين خاصية `background-image` لها.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.backgroundImage);
```

يجب أن يكون لديك محدد خاصية لاستهداف عناصر `span` التي تحمل الصنف `status` والتي هي عناصر تابعة لـ `tr` التي تحمل الصنف `read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]'));
```

يجب استخدام محدد خاصية لاستهداف عناصر `span` التي تحمل الصنف `status` والتي هي عناصر تابعة لـ `tr` التي تحمل الصنف `read` وتعيين خاصية `border` لها.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.border);
```

يجب استخدام محدد خاصية لاستهداف عناصر `span` التي تحمل الصنف `status` والتي هي عناصر تابعة لـ `tr` التي تحمل الصنف `read` وتعيين خاصية `background-image` لها.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.backgroundImage);
```

يجب أن يكون لديك محدد خاصية لاستهداف عناصر `span` التي تحمل الصنف `status` والتي هي عناصر تابعة لـ `tr` التي تحمل الصنف `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]'));
```

يجب استخدام محدد خاصية لاستهداف عناصر `span` التي تحمل الصنف `status` والتي هي عناصر تابعة لـ `tr` التي تحمل الصنف `in-progress` وتعيين خاصية `border` لها.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.border);
```

يجب استخدام محدد خاصية لاستهداف عناصر `span` التي تحمل الصنف `status` والتي هي عناصر تابعة لـ `tr` التي تحمل الصنف `in-progress` وتعيين خاصية `background-image` لها.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.backgroundImage);
```

يجب أن يكون لديك محدد خاصية لاستهداف عناصر `span` التي تحمل الصنف `status` وعناصر `span` التي تبدأ قيمة صنفها بـ `rate`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.exists(selector1 || selector2);
```

يجب استخدام محدد خاصية لاستهداف عناصر `span` التي تحمل الصنف `status` وعناصر `span` التي تبدأ قيمة صنفها بـ `rate` وتعيين خاصية `height` لها.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.height || selector2?.height);
```

يجب استخدام محدد خاصية لاستهداف عناصر `span` التي تحمل الصنف `status` وعناصر `span` التي تبدأ قيمة صنفها بـ `rate` وتعيين خاصية `width` لها.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.width || selector2?.width);
```

يجب استخدام محدد خاصية لاستهداف عناصر `span` التي تحمل الصنف `status` وعناصر `span` التي تبدأ قيمة صنفها بـ `rate` وتعيين خاصية `padding` لها.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.padding || selector2?.padding);
```

يجب أن يكون لديك محدد خاصية لاستهداف عناصر `span` التي هي أبناء مباشرين لعناصر `span` التي تبدأ قيمة `class` الخاصة بها بـ `rate`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span'));
```

يجب استخدام محدِّد خاصية لاستهداف عناصر `span` التي هي أبناء مباشرين لعناصر `span` التي تحتوي قيمة `class` تبدأ بـ `rate` وتعيين خاصية `border` لها.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border'));
```

يجب استخدام محدِّد خاصية لاستهداف عناصر `span` التي هي أبناء مباشرين لعناصر `span` التي تبدأ قيمة `class` الخاصة بها بـ `rate`، وضبط خاصية `border-radius` الخاصة بها.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border-radius'));
```

يجب استخدام محدِّد خاصية لاستهداف عناصر `span` التي هي أبناء مباشرين لعناصر `span` التي تبدأ قيمة الخاصية `class` فيها بـ `rate`، وضبط خاصية `margin` لها.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('margin'));
```

يجب استخدام محدِّد خاصية لاستهداف عناصر `span` التي هي أبناء مباشرين لعناصر `span` التي تبدأ قيمة الخاصية `class` فيها بـ `rate`، وضبط خاصية `height` لها.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('height'));
```

يجب استخدام محدد خاصية لاستهداف عناصر `span` التي هي أبناء مباشرين لعناصر `span` التي تبدأ قيمة `class` الخاصة بها بـ `rate`، وضبط خاصية `width` لها.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('width'));
```

يجب استخدام محدِّد خاصية لاستهداف عناصر `span` التي هي أبناء مباشرين لعناصر `span` التي تبدأ قيمة الخاصية `class` فيها بـ `rate`، وضبط خاصية `background-color` لها.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('background-color'));
```

يجب أن يكون لديك محدِّد خاصية لاستهداف أول عنصر تابع من عناصر `span` التي تحتوي على الكلمة `one` كجزء من قيمة `class` الخاصة بها.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

يجب استخدام محدِّد خاصية لاستهداف أول عنصر تابع من عناصر `span` التي تحتوي على الكلمة `one` كجزء من قيمة `class` الخاصة بها، وضبط خاصية `background-image` الخاصة به لاستخدام `linear-gradient`.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image').includes('linear-gradient('));
```

يجب أن يكون لديك محدِّد خصائص لاستهداف أول منحدرين من عناصر `span` التي تتضمن الكلمة `two` كجزء من قيمة `class` لها.

```js

const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

يجب استخدام محدِّد خاصية لاستهداف أول اثنين من العناصر التابعة لـ `span` التي تحتوي على الكلمة `two` كجزء من قيمة `class` الخاصة بها، وضبط خاصية `background-image` لديهم لاستخدام `linear-gradient`.

```js
const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]

assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.backgroundImage.includes('linear-gradient(')) 
```

يجب أن يكون لديك محدِّد خاصية لاستهداف عناصر `span` التي هي عناصر تابعة لعناصر `span` التي تحتوي على كلمة `three` كجزء من قيمة `class` الخاصة بها.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.exists(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

يجب استخدام محدِّد خاصية لاستهداف عناصر `span` التابعة لعناصر `span` التي تحتوي على كلمة `three` كجزء من قيمة `class` الخاصة بها، وضبط خاصية `background-image` لديهم لاستخدام `linear-gradient`.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.include(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image'), 'linear-gradient(');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main>
    <h1>Book Inventory</h1>
    <div class="table-container">
      <table>
        <thead>
          <tr>
            <th>Title</th>
            <th>Author</th>
            <th>Category</th>
            <th>Status</th>
            <th>Rate</th>
          </tr>
        </thead>
        <tbody>
          <tr class="read">
            <td>The Three Musketeers</td>
            <td>Alexandre Dumas</td>
            <td>Historical Novel</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate three">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="to-read">
            <td>The Plague</td>
            <td>Albert Camus</td>
            <td>Philosophical Novel</td>
            <td><span class="status">To Read</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>The Metamorphosis</td>
            <td>Franz Kafka</td>
            <td>Novella</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate one">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Dead Souls</td>
            <td>Nikolai Gogol</td>
            <td>Picaresque</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="in-progress">
            <td>Lord of the Flies</td>
            <td>William Golding</td>
            <td>Allegorical Novel</td>
            <td><span class="status">In Progress</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Do Androids Dream of Electric Sheep?</td>
            <td>Philip K. Dick</td>
            <td>Science Fiction</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
        </tbody>
      </table>
    </div>
  </main>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
    font-family: Arial, sans-serif;
}

.table-container {
  overflow-x: auto;
}

table {
    border-collapse: collapse;
    width: 100vw;
}

thead {
    background-image: linear-gradient(hsl(197, 92%, 77%, 0.60), hsl(197, 92%, 50%, 0.60));
}

th {
    padding: 0.4em;
}

td {
    text-align: center;
    padding: 0.3em;
}

tr[class="read"] {
    background-image: linear-gradient(hsl(120, 100%, 85%, 0.6), hsl(120, 100%, 65%, 0.6));
}

tr[class="to-read"] {
    background-image: linear-gradient(hsl(120, 20%, 95%, 0.6), hsl(120, 20%, 75%, 0.6));
}

tr[class="in-progress"] {
    background-image: linear-gradient(hsl(40, 100%, 85%, 0.60), hsl(40, 100%, 65%, 0.6));
}

span {
    display: inline-block;
}

span[class="status"] {
    border-radius: 20%/60%;
}

span[class="status"],
span[class^="rate"] {
    height: 1.85em;
    width: 7em;
    padding: 0.3em;
}

tr[class="to-read"] span[class="status"] {
    border: 0.1em solid hsl(5, 100%, 45%);
    background-image: linear-gradient(hsl(5, 100%, 75%), hsl(5, 100%, 50%));
}

tr[class="read"] span[class="status"] {
    border: 0.1em solid hsl(120, 75%, 45%);
    background-image: linear-gradient(hsl(120, 75%, 75%), hsl(120, 75%, 50%));
}

tr[class="in-progress"] span[class="status"] {
    border: 0.1em solid hsl(40, 90%, 40%);
    background-image: linear-gradient(hsl(40, 90%, 75%), hsl(40, 90%, 50%));
}

span[class^="rate"] > span {
    border: 0.1em solid hsl(0, 0%, 50%);
    border-radius: 50%;
    background-color: hsl(0, 15%, 95%);
    height: 100%;
    width: 20%;
    margin: 0.1em;
}

span[class~="three"] span {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="two"] :nth-child(1),
span[class~="two"] :nth-child(2) {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="one"] :first-child {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}
```
