---
id: 587d78b2367417b2b2512b10
title: 使用 splice() 删除元素
challengeType: 1
forumTopicId: 301166
dashedName: remove-items-using-splice
---

# --description--

好了，我们已经学会了如何使用 `shift()` 和 `pop()` 从数组的开头和结尾移除元素，但如果我们想从中间某个位置移除元素呢？或者一次移除多个元素呢？这时就用到 `splice()` 了。`splice()` 允许我们做到这一点：**从数组的任意位置移除任意数量的连续元素**。

`splice()` 最多可以接受 3 个参数，但现在我们只关注前两个。`splice()` 的前两个参数是整数，表示调用 `splice()` 的数组中项的索引或位置。记住，数组是*零索引的*，所以表示数组第一个元素时，我们用 `0`。`splice()` 的第一个参数表示从数组中哪个索引开始移除元素，第二个参数表示要删除的元素数量。例如：

```js
let array = ['today', 'was', 'not', 'so', 'great'];

array.splice(2, 2);
```

这里我们移除 2 个元素，从第三个元素（索引 2）开始。`array` 的值将是 `['today', 'was', 'great']`。

`splice()` 不仅会修改调用该方法的数组，还会返回一个包含被移除元素的数组：

```js
let array = ['I', 'am', 'feeling', 'really', 'happy'];

let newArray = array.splice(3, 2);
```

`newArray` 值为 `['really', 'happy']`。

# --instructions--

我们已经定义了数组 `arr`。 请使用 `splice()` 从 `arr` 里移除元素，使剩余的元素之和为 `10`。

# --hints--

不应修改这一行 `const arr = [2, 4, 5, 1, 7, 5, 2, 1];`。

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/constarr=\[2,4,5,1,7,5,2,1\];?/)
);
```

`arr` 的剩余元素之和应为 `10`。

```js
assert.strictEqual(
  arr.reduce((a, b) => a + b),
  10
);
```

应对 `arr` 调用 `splice()` 方法。

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/arr\.splice\(/));
```

splice 应只删除 `arr` 里面的元素，不能给 `arr` 添加元素。

```js
assert(
  !__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/arr\.splice\(\d+,\d+,\d+.*\)/g)
);
```

# --seed--

## --seed-contents--

```js
const arr = [2, 4, 5, 1, 7, 5, 2, 1];
// Only change code below this line

// Only change code above this line
console.log(arr);
```

# --solutions--

```js
const arr = [2, 4, 5, 1, 7, 5, 2, 1];
arr.splice(1, 4);
```
