---
id: 68e7e729f6d14f69a23786ef
title: Hatua ya 2
challengeType: 0
dashedName: step-2
---

# --description--

Ifuatayo, ndani ya kipengele chako cha `main`, ongeza kipengele cha `h1` chenye maandishi ya `Parent Teacher Conference Form` na darasa la vitu `title` na `center`. 

# --hints--

Unapaswa kuwa na kipengele cha `h1`.

```js
assert.exists(document.querySelector("h1"));
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi ya `Parent Teacher Conference Form`.

```js
assert.strictEqual(document.querySelector("h1")?.innerText, "Parent Teacher Conference Form");
```

Kipengele chako cha `h1` kinapaswa kuwa na madarasa ya vitu `title` na `center`.

```js
assert.exists(document.querySelector("h1.title.center"));
```

Kipengele chako cha `h1` kinapaswa kuwa ndani ya kipengele cha `main`.

```js
assert.exists(document.querySelector("main h1.title.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
--fcc-editable-region--
    
--fcc-editable-region--
  </main>
</body>

</html>
```
