---
id: 68e707257ea91c2e12d489fb
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utafundishwa jinsi ya kuongeza mitindo maalum kwa vitufe vya radio kwa kujenga fomu ya mkutano wa wazazi na walimu. Muundo wa msingi wa HTML umeandaliwa tayari kwa ajili yako.

Anza kwa kuongeza kipengele cha `main` chenye darasa la vitu linaloitwa `container`.

# --hints--

Unapaswa kuwa na kipengele cha `main`.

```js
assert.exists(document.querySelector("main"));
```

Kipengele chako cha `main` kinapaswa kuwa na darasa la vitu linaloitwa `container`.

```js
assert.exists(document.querySelector("main.container"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
