---
id: 68eaa668d4d9fbe529a0e16b
title: Hatua 14
challengeType: 0
dashedName: step-14
---

# --description--

Sasa, ongeza kipengele cha `label` chenye maandishi ya `Any specific concerns or topics you'd like to discuss?` na sifa ya `for` iliyowekwa kuwa `notes`.

Chini ya `label` hii, ongeza kipengele cha `textarea`. Weka `id` yake kuwa `notes`, `name` kuwa `notes`, `rows` kuwa `4`, na `cols` kuwa `50`.

# --hints--

`fieldset` yako ya mwisho inapaswa kuwa na kipengele cha `label`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

Kipengele cha `label` kinapaswa kuwa na maandishi ya `Any specific concerns or topics you'd like to discuss?`.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.innerText, "Any specific concerns or topics you'd like to discuss?");
```

Sifa ya `label` ya kipengele cha `for` inapaswa kuwekwa kuwa `"notes"`.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "notes");
```

`fieldset` yako ya mwisho inapaswa kuwa na kipengele cha `textarea`.

```js
assert.exists(document.querySelector("fieldset:last-of-type textarea"));
```

Kipengele chako cha `textarea` kinapaswa kuwa chini ya kipengele cha `label`.

```js
const label = document.querySelector("fieldset:last-of-type label");
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(label?.nextElementSibling, textarea);
```

Kipengele cha `textarea` kinapaswa kuwa na sifa ya `id` iliyowekwa kuwa `"notes"`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("id"), "notes");
```

Kipengele cha `textarea` kinapaswa kuwa na sifa ya `name` iliyowekwa kuwa `"notes"`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("name"), "notes");
```

Kipengele cha `textarea` kinapaswa kuwa na sifa ya `rows` iliyowekwa kuwa `4`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("rows"), "4");
```

Kipengele cha `textarea` kinapaswa kuwa na sifa ya `cols` iliyowekwa kuwa `50`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("cols"), "50");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
