---
id: 68ea975c7cbbb894a386e402
title: Hatua 11
challengeType: 0
dashedName: step-11
---

# --description--

Hatua inayofuata ni kuongeza vipengele vya `label` na `input` kwa njia ya mawasiliano ya barua pepe.

Anza kwa kuongeza lebo yenye darasa la `"contact-method"` na sifa ya `for` iliyowekwa kuwa `"email"`. Maandishi ya lebo yanapaswa kuwa `Email: `.

Chini ya kipengele chako cha `label`, ongeza kitufe cha radio chenye sifa za `id` na `value` zilizowekwa kuwa `"email"`. Sifa ya `name` inapaswa kuwekwa kuwa `"contact-method"` na darasa liwe `"contact-method-radio-btn"`.

Mwishowe, hakikisha kitufe hiki cha radio kimekuwa checked kwa chaguo-msingi.

# --hints--


Seti yako ya sehemu ya mwisho inapaswa kuwa na kipengele cha `label`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

Kipengele hiki kipya cha `label` (kwa barua pepe) kinapaswa kuwa na darasa la `"contact-method"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

Kipengele kipya cha `label` kina sifa ya `for` inayopaswa kuwekwa kuwa `"email"`.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "email");
```

Kipengele cha `label` kinapaswa kuwa na maandishi ya `Email: `. Hakikisha kuna nafasi baada ya nukta.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.match(label?.innerText, /Email:\s?/);
```

Unapaswa kuwa na kipengele cha `input` ndani ya `fieldset` ya mwisho.

```js
assert.exists(document.querySelector("fieldset:last-of-type input"));
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya `type` iliyowekwa kuwa `radio`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

Kipengele cha `input` kinapaswa kuwa na sifa ya `id` iliyowekwa kuwa `"email"`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("id"), "email");
```

Kipengele cha `input` kinapaswa kuwa na sifa ya `value` iliyowekwa kuwa `"email"`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("value"), "email");
```

Kipengele cha `input` kinapaswa kuwa na sifa ya `name` iliyowekwa kuwa `"contact-method"`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

Kipengele cha `input` kinapaswa kuwa na darasa la `"contact-method-radio-btn"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type input.contact-method-radio-btn"));
```

Kitufe chako cha `radio` kinapaswa kuwa `checked` kwa chaguo-msingi.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.isTrue(input?.hasAttribute("checked"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
