---
id: 68eafe1bae175f3b2a87a90d
title: Hatua 15
challengeType: 0
dashedName: step-15
---

# --description--

Mwishowe, unapaswa kuongeza kitufe cha kuwasilisha kwa fomu.

Chini ya kipengele cha mwisho cha `fieldset`, ongeza kipengele cha `button` chenye darasa la `submit-btn` na sifa ya kipengele ya `type` iliyowekwa kuwa `"submit"`. Maandishi ya kitufe yanapaswa kuwa `Submit Form`.

Kwa hili, muundo wako wa HTML umekamilika!

# --hints--

Unapaswa kuongeza kipengele cha `button` chini ya kipengele cha mwisho cha `fieldset`.

```js
assert.exists(document.querySelector("fieldset:last-of-type + button"));
```

Kipengele cha `button` kinapaswa kuwa na darasa la `"submit-btn"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type + button.submit-btn"));
```

Kipengele cha `button` kinapaswa kuwa na sifa ya kipengele ya `type` iliyowekwa kuwa `"submit"`.

```js
const button = document.querySelector("fieldset:last-of-type + button");
assert.strictEqual(button?.getAttribute("type"), "submit");
```

Kipengele cha `button` kinapaswa kuwa na maandishi ya `Submit Form`.

```js
const button = document.querySelector("fieldset:last-of-type + button");
assert.strictEqual(button?.innerText, "Submit Form");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
