---
id: 68ea7c41f2555e16174e7cfa
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Ifuatayo, ongeza kipengele cha aya chenye maandishi `Please fill out the form below to help schedule your parent-teacher conference.`. Kipengele chako cha aya kinapaswa pia kuwa na madarasa ya vitu `description` na `center`.

# --hints--

Unapaswa kuwa na kipengele cha `p`.

```js
assert.exists(document.querySelector("p"));
```

Kipengele chako cha `p` kinapaswa kuwa na maandishi `Please fill out the form below to help schedule your parent-teacher conference.`

```js
assert.strictEqual(document.querySelector("p")?.innerText, "Please fill out the form below to help schedule your parent-teacher conference.");
```

Kipengele chako cha `p` kinapaswa kuwa na madarasa ya vitu `description` na `center`.

```js
assert.exists(document.querySelector("p.description.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </main>
</body>

</html>
```
