---
id: 68ea95c833f0518ae6e366cd
title: Hatua 10
challengeType: 0
dashedName: step-10
---

# --description--

Katika hatua chache zijazo, utaongeza vipengele vya fomu vinavyohusika na kukusanya njia ya mawasiliano inayopendelewa na mtumizi.

Anza kwa kuongeza kipengele kingine cha `fieldset` chenye kipengele cha `legend` kilichopangwa ndani. Kipengele chako cha `legend` kinapaswa kuwa na maandishi ya `Preferred Contact Method`.

# --hints--

Una jumla ya vipengele vitatu vya `fieldset` kwenye ukurasa.

```js
assert.lengthOf(document.querySelectorAll("fieldset"), 3);
```

Unapaswa kuwa na kipengele cha `legend` ndani ya kipengele chako cha tatu cha `fieldset`.

```js
assert.exists(document.querySelector("fieldset:last-of-type legend"));
```

Kipengele chako cha `legend` kinapaswa kuwa na maandishi ya `Preferred Contact Method`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type legend")?.innerText, "Preferred Contact Method");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
