---
id: 68ea8c330ed07b60879f3d38
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Sasa ni wakati wa kuongeza vipengele vya fomu ili kukusanya taarifa za alama za mwanafunzi.

Anza kwa kuongeza kipengele kingine cha `label` chenye maandishi ya `Student Grade: ` na sifa ya `for` kuwekwa kuwa `"grade"`.

Kisha, chini ya kipengele chako cha `label`, ongeza kipengele cha `input` chenye sifa ya `type` kuwekwa kuwa `"number"`. Sifa za `name` na `id` zinapaswa kuwekwa kuwa `"grade"`. Sifa ya `placeholder` inapaswa kuwekwa kuwa `"E.g., 4"`. Mwishowe, nambari yako ya `input` inapaswa kuwa ya lazima. 

# --hints--

Unapaswa kuwa na jumla ya vipengele viwili vya `label` kwenye ukurasa.

```js
assert.lengthOf(document.querySelectorAll("label"), 2);
```

Kipengele chako cha pili cha `label` kinapaswa kuwa na maandishi ya `Student Grade: `. Usisahau nafasi baada ya nukta mbili.

```js
assert.match(code, /Student\s+Grade:\s+<\/label>/);
```

Kipengele chako cha pili cha `label` kinapaswa kuwa na sifa ya `for`.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

Sifa ya `label` ya kipengele chako cha pili cha `for` inapaswa kuwa na thamani ya `"grade"`.

```js
assert.strictEqual(document.querySelector("label:last-of-type")?.getAttribute("for"), "grade");
```

Unapaswa kuwa na vipengele viwili vya `input` kwenye ukurasa.

```js
assert.lengthOf(document.querySelectorAll("input"), 2);
```

Kipengele chako cha pili cha `input` kinapaswa kuwa chini ya kipengele chako cha pili cha `label`.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const inputEl = document.querySelector("fieldset input:last-of-type");
const labelEl = document.querySelector("label:last-of-type");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

Kipengele chako cha pili cha `input` kinapaswa kuwa na sifa ya `type` iliyowekwa kuwa `"number"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("type"), "number");
```

Kipengele chako cha pili cha `input` kinapaswa kuwa na sifa ya `name` iliyowekwa kuwa `"grade"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("name"), "grade");
```

Kipengele chako cha pili cha `input` kinapaswa kuwa na sifa ya `id` iliyowekwa kuwa `"grade"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("id"), "grade");
```

Kipengele chako cha pili cha `input` kinapaswa kuwa na sifa ya `placeholder` iliyowekwa kuwa `"E.g., 4"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., 4");
```

Kipengele chako cha pili cha `input` kinapaswa kuwa na sifa ya `required`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
