---
id: 68ea8e6a87a2816b0bb2eb7d
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Hatua inayofuata ni sehemu katika fomu kwa ajili ya taarifa za mzazi.

Anza kwa kuongeza kipengele kingine cha `fieldset`. Ndani ya kipengele hicho cha `fieldset`, ongeza kipengele cha `legend` chenye maandishi ya `Parent/Guardian Information`.

# --hints--

Una jumla ya vipengele viwili vya `fieldset` kwenye ukurasa.

```js
assert.lengthOf(document.querySelectorAll("fieldset"), 2);
```

Unapaswa kuwa na kipengele cha `legend` ndani ya kipengele chako cha pili cha `fieldset`.

```js
assert.exists(document.querySelector("fieldset:last-of-type legend"));
```

Kipengele chako cha `legend` kinapaswa kuwa na maandishi ya `Parent/Guardian Information`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type legend")?.innerText, "Parent/Guardian Information");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
