---
id: 68ea9f01c51b66b937793611
title: Hatua 12
challengeType: 0
dashedName: step-12
---

# --description--

Ifuatayo, ongeza lebo nyingine yenye darasa la `"contact-method"` na sifa ya kipengele `for` iliyowekwa kuwa `"phone"`. Maandishi ya lebo yanapaswa kuwa `Phone: `.

Chini ya kipengele chako cha `label`, ongeza kitufe cha radio chenye sifa za `id` na `value` zilizowekwa kuwa `"phone"`. Sifa ya `name` inapaswa kuwekwa kuwa `"contact-method"` na darasa liwe `"contact-method-radio-btn"`.

# --hints--

Unapaswa kuwa na vipengele viwili vya `label` ndani ya `fieldset` ya mwisho.

```js
assert.lengthOf(document.querySelectorAll("fieldset:last-of-type label"), 2);
```

Kipengele cha `label` kinapaswa kuwa na darasa la `"contact-method"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

Sifa ya kipengele cha `label` ya `for` inapaswa kuwekwa kuwa `"phone"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label[for='phone']"));
```

Kipengele cha `label` kinapaswa kuwa na maandishi ya `Phone: `.

```js
const label = document.querySelector("fieldset:last-of-type label[for='phone']");
assert.match(label?.innerText, /Phone:\s?/);
```

Unapaswa kuwa na vipengele viwili vya `input` ndani ya `fieldset` ya mwisho.

```js
assert.lengthOf(document.querySelectorAll("fieldset:last-of-type input"), 2);
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele `type` iliyowekwa kuwa `radio`.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

Kipengele cha `input` kinapaswa kuwa na sifa ya kipengele `id` iliyowekwa kuwa `"phone"`.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("id"), "phone");
```

Kipengele cha `input` kinapaswa kuwa na sifa ya kipengele `value` iliyowekwa kuwa `"phone"`.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("value"), "phone");
```

Kipengele cha `input` kinapaswa kuwa na sifa ya kipengele `name` iliyowekwa kuwa `"contact-method"`.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

Kipengele cha `input` kinapaswa kuwa na darasa la `"contact-method-radio-btn"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type input[id='phone'].contact-method-radio-btn"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
