---
id: 68ea8a5db73b155708992687
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Hatua inayofuata ni kuongeza kipengele cha `input` kinachohusiana na taarifa za mwanafunzi.

Anza kwa kuongeza kipengele cha `input` chenye sifa ya kipengele ya `type` iliyowekwa kuwa `"text"`. Kisha ongeza sifa ya kipengele ya `name` na `id`, zote zikiwa zimewekwa kuwa `student-name`.

Ifuatayo, ongeza sifa ya kipengele ya `placeholder` iliyowekwa kuwa `E.g., Jane Doe`. Na hatimaye, ongeza sifa ya kipengele ya `required`.

# --hints--

Unapaswa kuwa na kipengele cha `input`.

```js
assert.exists(document.querySelector("fieldset input"));
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `type` iliyowekwa kuwa `"text"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `name` iliyowekwa kuwa `"student-name"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("name"), "student-name");
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `id` iliyowekwa kuwa `"student-name"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("id"), "student-name");
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `placeholder` iliyowekwa kuwa `E.g., Jane Doe`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Jane Doe");
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `required`.

```js
const inputEl = document.querySelector("fieldset input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>

      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>


    </form>
  </main>
</body>

</html>
```
