---
id: 68ea8666b8b473455f5b33ef
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Sehemu ya kwanza ya fomu itazingatia taarifa za mwanafunzi, kama jina lao na daraja.

Ndani ya kipengele cha `form`, ongeza kipengele cha `fieldset`. Ndani ya kipengele hicho cha `fieldset`, ongeza kipengele cha `legend` chenye maandishi ya `Student Information`. 

Chini ya kipengele cha `legend`, ongeza kipengele cha `label` chenye maandishi ya `Full Name: ` na sifa ya `for` yenye thamani ya `student-name`.


# --hints--

Unapaswa kuwa na kipengele cha `fieldset`.

```js
assert.exists(document.querySelector("fieldset"));
```

Kipengele chako cha `fieldset` kinapaswa kuwa ndani ya kipengele cha `form`.

```js
assert.exists(document.querySelector("form fieldset"));
```

Unapaswa kuwa na kipengele cha `legend`.

```js
assert.exists(document.querySelector("legend"));
```

Kipengele chako cha `legend` kinapaswa kuwa ndani ya kipengele cha `fieldset`.

```js
assert.exists(document.querySelector("fieldset legend"));
```

Kipengele chako cha `legend` kinapaswa kuwa na maandishi ya `Student Information`.

```js
assert.strictEqual(document.querySelector("fieldset legend")?.innerText, "Student Information");
```

Unapaswa kuwa na kipengele cha `label`.

```js
assert.exists(document.querySelector("label"));
```

Kipengele chako cha `label` kinapaswa kuwa na maandishi ya `Full Name: `. Usisahau nafasi baada ya nukta mbili.

```js
assert.match(code, /Full\s+Name:\s+<\/label>/);
```

Kipengele chako cha `label` kinapaswa kuwa na sifa ya `for`.

```js
assert.isTrue(document.querySelector("label")?.hasAttribute("for"));
```

Sifa ya `label` ya kipengele chako cha `for` inapaswa kuwa na thamani ya `student-name`.

```js
assert.strictEqual(document.querySelector("label")?.getAttribute("for"), "student-name");
```

Kipengele chako cha `label` kinapaswa kuwa chini ya kipengele cha `legend`.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const legendEl = document.querySelector("fieldset legend");
const labelEl = document.querySelector("label");

assert.isBelow(fieldsetChildren.indexOf(legendEl), fieldsetChildren.indexOf(labelEl));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
