---
id: 68ea8fdf2c9b167307e7f671
title: Hatua 9
challengeType: 0
dashedName: step-9
---

# --description--

Sasa ni wakati wa kuongeza vipengele vya fomu ili kukusanya taarifa za mzazi.

Anza kwa kuongeza kipengele cha `label` chenye maandishi ya `Parent/Guardian Name: ` na sifa ya kipengele ya `for` ikiwa imewekwa kuwa `"parent-name"`.

Kisha, chini ya kipengele chako cha `label`, ongeza kipengele cha `input` chenye sifa ya kipengele ya `type` ikiwa imewekwa kuwa `"text"`. Sifa za kipengele za `name` na `id` zinapaswa kuwekwa kuwa `"parent-name"`. Sifa ya kipengele ya `placeholder` inapaswa kuwekwa kuwa `"E.g., Nancy Doe"`. Mwishowe, kipengele chako cha `input` kinapaswa kuwa kinahitajika.

# --hints--

Unapaswa kuwa na jumla ya vipengele vitatu vya `label` kwenye ukurasa.

```js
assert.lengthOf(document.querySelectorAll("label"), 3);
```

Kipengele chako cha tatu cha `label` kinapaswa kuwa na maandishi ya `Parent/Guardian Name: `. Usisahau nafasi baada ya nukta mbili.

```js
assert.match(code, /Parent\/Guardian\s+Name:\s+<\/label>/);
```

Kipengele chako cha tatu cha `label` kinapaswa kuwa na sifa ya kipengele ya `for`.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

Sifa ya kipengele ya `label` ya kipengele chako cha tatu cha `for` inapaswa kuwa na thamani ya `"parent-name"`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type label")?.getAttribute("for"), "parent-name");
```

Unapaswa kuwa na vipengele vitatu vya `input` kwenye ukurasa.

```js
assert.lengthOf(document.querySelectorAll("input"), 3);
```

Kipengele chako cha tatu cha `input` kinapaswa kuwa chini ya kipengele chako cha tatu cha `label`.

```js
const fieldsetChildren = [...document.querySelector("fieldset:last-of-type")?.children];
const inputEl = document.querySelector("fieldset:last-of-type input");
const labelEl = document.querySelector("fieldset:last-of-type label");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

Kipengele chako cha tatu cha `input` kinapaswa kuwa na sifa ya kipengele ya `type` ikiwa imewekwa kuwa `"text"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

Kipengele chako cha tatu cha `input` kinapaswa kuwa na sifa ya kipengele ya `name` ikiwa imewekwa kuwa `"parent-name"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("name"), "parent-name");
```

Kipengele chako cha tatu cha `input` kinapaswa kuwa na sifa ya kipengele ya `id` ikiwa imewekwa kuwa `"parent-name"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("id"), "parent-name");
```

Kipengele chako cha tatu cha `input` kinapaswa kuwa na sifa ya kipengele ya `placeholder` ikiwa imewekwa kuwa `"E.g., Nancy Doe"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Nancy Doe");
```

Kipengele chako cha tatu cha `input` kinapaswa kuwa na sifa ya kipengele ya `required`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
