---
id: 68eb058081504963447bfa60
title: Hatua 17
challengeType: 0
dashedName: step-17
---

# --description--

Ifuatayo, unapaswa kupamba kontena kuu. Chagua kipengele chenye darasa la vitu `container` na weka `background-color` yake kuwa msimbo wa rangi wa hex `#ffffff1a`.

Msimbo wa hex ni mchanganyiko wa tarakimu sita za nambari na herufi zinazotumika katika HTML na CSS kuwakilisha rangi. Huanza na `#` ikifuatiwa na jozi tatu:

- Tarakimu mbili za kwanza zinawakilisha rangi nyekundu
- Tarakimu mbili zinazofuata zinawakilisha rangi ya kijani
- Tarakimu mbili za mwisho zinawakilisha rangi ya buluu

Kwa mfano, `#ffffff` ni nyeupe safi kwa sababu ina thamani kubwa kabisa kwa nyekundu, kijani, na buluu.

Pia unaweza kuongeza tarakimu mbili za ziada mwishoni kudhibiti uwazi (inayoitwa chaneli ya alfa). Katika `#ffffff1a`, `1a` hufanya rangi nyeupe kuwa na uwazi wa nusu. Thamani ya alfa ikiwa ndogo, rangi inaonekana kuwa wazi zaidi. Utajifunza zaidi kuhusu misimbo ya hex baadaye!

# --hints--

Unapaswa kulenga kipengele chenye darasa la vitu `container`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

Weka sifa ya `background-color` kuwa `#ffffff1a`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.container')
assert.strictEqual(style?.backgroundColor, 'rgba(255, 255, 255, 0.1)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```


```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

--fcc-editable-region--

--fcc-editable-region--
```
