---
id: 587d7fb0367417b2b2512bef
title: 提供 HTML 檔案
challengeType: 2
forumTopicId: 301516
dashedName: serve-an-html-file
---

# --description--

你可以使用 `res.sendFile(path)` 方法以檔案回應請求。你可以將它放在 `app.get('/', ...)` 路由處理常式中。在背後，這個方法會設定適當的標頭，以指示你的瀏覽器如何根據檔案的型別來處理你想要傳送的檔案。接著它會讀取並傳送該檔案。此方法需要絕對檔案路徑。我們建議你使用 Node 全域變數 `__dirname` 來計算路徑，如下所示：

```js
absolutePath = __dirname + '/relativePath/file.ext'
```

# --instructions--

將 `/views/index.html` 檔案作為對 `/` 路徑的 GET 請求的回應傳送。如果你查看你的即時應用程式，你應該會看到一個大型的 HTML 標題（以及我們稍後會使用的表單……），且沒有套用任何樣式。

**注意：** 你可以編輯前一個挑戰的解決方案或創建一個新的解決方案。如果你創建新的解決方案，請記住 Express 是從上到下評估路由，並執行第一個符合條件的處理常式。你必須將前一個解決方案註解掉，否則伺服端會持續回應字串。

# --hints--

你的應用程式應該提供檔案 views/index.html

```js
  const response = await fetch(code);
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  assert.match(
    data,
    /<h1>.*<\/h1>/,
    'Your app does not serve the expected HTML'
  );
```

