---
id: 587d7fb1367417b2b2512bf3
title: 實作一個根層級請求記錄中介層
challengeType: 2
forumTopicId: 301514
dashedName: implement-a-root-level-request-logger-middleware
---

# --description--

之前，你已經介紹過 `express.static()` 中介層函式（程式）。現在是時候更詳細地了解中介層是什麼。中介層函式是接受 3 個引數的函式：請求物件、回應物件，以及應用程式請求-回應週期中的下一個函式。這些函式會執行一些可能對應用程式產生副作用的程式碼，並且通常會為請求或回應物件添加資訊。當某些條件符合時，它們也可以透過傳送回應來結束週期。如果它們完成時沒有傳送回應，則會開始執行堆疊中的下一個函式。這會觸發呼叫第三個引數，`next()`。

請看以下範例：

```js
function(req, res, next) {
  console.log("I'm a middleware...");
  next();
}
```

假設你將這個函式掛載在一個路由上。當請求符合該路由時，它會顯示字串 “I’m a middleware…” ，然後執行堆疊中的下一個函式。在這個練習中，你將建置根層級的中介層。正如你在挑戰 4 中看到的，要在根層級掛載中介層函式，你可以使用 `app.use(<mware-function>)` 方法。在這種情況下，該函式會對所有請求執行，但你也可以設定更具體的條件。舉例來說，如果你想讓某個函式只對 POST 請求執行，你可以使用 `app.post(<mware-function>)`。所有 HTTP 動詞（GET、DELETE、PUT、…）都有類似的方法。

# --instructions--

建置一個簡單的記錄器。對於每個請求，它應該在主控台記錄一個字串，格式如下：`method path - ip`。範例如下：`GET /json - ::ffff:127.0.0.1`。請注意 `method` 和 `path` 之間有一個空格，且分隔 `path` 和 `ip` 的破折號兩邊也各有一個空格。你可以從請求物件使用 `req.method`、`req.path` 和 `req.ip` 取得請求方法（http 動詞）、相對路由路徑和呼叫者的 ip。完成後記得呼叫 `next()`，否則你的伺服器將會永遠卡住。請確保已開啟「Logs」，並觀察當有請求到達時會發生什麼。

**注意：** Express 會依照它們在程式碼中出現的順序評估函式（程式）。中介層也是如此。如果你想讓它對所有路由生效，應該將它掛載在路由之前。

# --hints--

根層級的 logger 中介層應該是作用中的

```js
  const response = await fetch(code + '/_api/root-middleware-logger');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(
    data.passed,
    'root-level logger is not working as expected'
  );
```

