---
id: 587d7fb0367417b2b2512bee
title: 啟動一個運作中的 Express 伺服器
challengeType: 2
forumTopicId: 301519
dashedName: start-a-working-express-server
---

# --description--

在檔案 `myApp.js` 的前兩行中，你可以看到創建 Express app 物件是多麼簡單。這個物件有多個方法，你將在這些挑戰中學習許多方法。一個基本的方法是 `app.listen(port)`。它告訴你的伺服器監聽指定的埠，讓它進入執行狀態。出於測試原因，我們需要 app 在背景中執行，因此我們已經在 `server.js` 檔案中為你添加了這個方法。

讓我們提供第一個字串！在 Express 中，路由採用以下結構：`app.METHOD(PATH, HANDLER)`。METHOD 是小寫的 http 方法。PATH 是伺服端上的相對路徑（它可以是字串，甚至是正規表達式）。HANDLER 是當路由匹配時 Express 呼叫的函式。處理常式採用 `function(req, res) {...}` 的形式，其中 req 是請求物件，res 是回應物件。例如，該處理常式

```js
function(req, res) {
  res.send('Response String');
}
```

將字串 'Response String' 提供服務。

# --instructions--

使用 `app.get()` 方法來回應符合 `/`（根目錄）路徑的 GET 請求，並傳送字串 "Hello Express"。請透過查看日誌來確保你的程式碼能正常運作。

**注意：** 這些課程的所有程式碼都應該添加在我們為你開始的幾行程式碼之間。

# --hints--

你的應用程式應該提供字串 'Hello Express'

```js
  const response = await fetch(code);
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  assert.equal(
    data,
    'Hello Express',
    'Your app does not serve the text "Hello Express"'
  );
```

