---
id: 587d7fb0367417b2b2512bf0
title: 提供靜態資源
challengeType: 2
forumTopicId: 301518
dashedName: serve-static-assets
---

# --description--

HTML 伺服端通常會有一個或多個使用者可以存取的目錄。你可以將應用程式所需的靜態資源（樣式表、腳本、影像）放置在那裡。

在 Express 中，你可以使用中介層 `express.static(path)` 來實現此功能，其中 `path` 參數是包含資源的資料夾的絕對路徑。

如果你不知道什麼是中介層……別擔心，我們稍後會詳細討論。基本上，中介層是攔截路由處理常式的函式，添加某種資訊。中介層需要使用方法 `app.use(path, middlewareFunction)` 來掛載。第一個 `path` 引數是選用的。如果你不傳入它，中介層將會對所有請求執行。

# --instructions--

將 `express.static()` 中介層掛載到路徑 `/public`，使用 `app.use()`。資產資料夾的絕對路徑是 `__dirname + /public`。

現在你的應用程式應該能夠提供 CSS 樣式表。請注意，專案樣板中的 `/views/index.html` 有取用 `/public/style.css` 檔案。你的首頁現在應該看起來更好一些！

# --hints--

你的應用程式應該從 `/public` 目錄提供資產檔案到 `/public` 路徑。

```js
  const response = await fetch(code + '/public/style.css');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  assert.match(
    data,
    /body\s*\{[^\}]*\}/,
    'Your app does not serve static assets'
  );
```

你的應用程式不應該從除 `/public` 目錄以外的任何其他資料夾提供檔案。 

```js
  const response = await fetch(code + '/server.js');
  if (response.ok) {
    const data = await response.text();
    assert.equal(
      response.status + '',
      404 + '',
      'Your app must serve files only from "public" directory'
    );
  } else {
    assert.equal(
      response.status + '',
      404 + '',
      'Your app must serve files only from "public" directory'
    );
  }
```

