---
id: 587d7fb1367417b2b2512bf2
title: 使用 `.env` 檔案
challengeType: 2
forumTopicId: 301521
dashedName: use-the--env-file
---

# --description--

`.env` 檔案是一個隱藏檔案，用來將環境變數傳遞給你的應用程式。這個檔案是秘密的，除了你之外沒有人可以存取，並且可以用來儲存你想要保密或隱藏的資料。 例如，你可以儲存來自外部服務的 API 金鑰或你的資料庫 URI。你也可以用它來儲存設定選項。透過設定設定選項，你可以改變應用程式的行為，而不需要重寫部分程式碼。

環境變數可以從應用程式中透過 `process.env.VAR_NAME` 存取。`process.env` 物件是全域的 Node 物件，變數皆以字串形式傳遞。依慣例，變數名稱皆為大寫，字與字之間以底線分隔。`.env` 是一個 shell 檔案，因此你不需要在名稱或值周圍加上引號。同時也要注意，當你為變數指定值時，等號兩邊不能有空格，例如 `VAR_NAME=value`。通常，你會將每個變數定義放在獨立的一行。

# --instructions--

讓我們將環境變數添加為設定選項。

在你的專案目錄根目錄中創建一個 `.env` 檔案，並將變數 `MESSAGE_STYLE=uppercase` 儲存在其中。

然後，在你於上一個挑戰中創建的 `/json` GET 路由處理常式中，存取 `process.env.MESSAGE_STYLE`，如果該變數相等於 `uppercase`，則將回應物件的 `message` 轉換為大寫。回應物件應該是 `{"message": "Hello json"}` 或 `{"message": "HELLO JSON"}`，取決於 `MESSAGE_STYLE` 的值。請注意，你必須在路由處理常式**內部**讀取 `process.env.MESSAGE_STYLE` 的值，而非在外部，這是因為我們的測試執行方式所致。

你需要使用 `dotenv` 軟體包。它會將環境變數從你的 `.env` 檔案載入到 `process.env`。`dotenv` 軟體包已經安裝，並且在你的專案的 `package.json` 檔案中。在你的 `myApp.js` 檔案頂部，添加 `require('dotenv').config()` 以載入環境變數。

# --hints--

端點 `/json` 的回應應根據環境變數 `MESSAGE_STYLE` 變更。

```js
  const response = await fetch(code + '/_api/use-env-vars');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(
    data.passed,
    'The response of "/json" does not change according to MESSAGE_STYLE'
  );
```

