---
id: 587d7fb0367417b2b2512bed
title: 認識 Node 的控制檯
challengeType: 2
forumTopicId: 301515
dashedName: meet-the-node-console
---

# --description--

Working on these challenges will involve you writing your code using one of the following methods:

- 克隆 <a href="https://github.com/freeCodeCamp/boilerplate-express/" target="_blank" rel="noopener noreferrer nofollow">GitHub 倉庫</a> 並完成您的本地項目。
- Use a site builder of your choice to complete the project. Be sure to incorporate all the files from our GitHub repo.

During the development process, it is important to be able to check what’s going on in your code.

Node is just a JavaScript environment. Like client side JavaScript, you can use the console to display useful debug information. On your local machine, you would see console output in a terminal.

We recommend to keep the terminal open while working at these challenges. By reading the output in the terminal, you can see any errors that may occur.

更改文件後必須重新啓動服務器。

你可以使用 `Ctrl + C` 從終端停止服務器，然後直接使用 Node（`node mainEntryFile.js`）或使用 `package.json` 文件中的運行腳本 `npm run` 啓動服務器。

例如，可以使用 `npm run start` 在終端運行 `"start": "node server.js"` 腳本。

要實現保存文件時服務器自動重啓，你可以將 Node 的 `--watch` 標誌添加到你的啓動腳本 `"start": "node --watch server.js"`，或者你可以安裝一個 npm 包，例如 `nodemon`。 我們將把這作爲一項練習留給你。

# --instructions--

Modify the `myApp.js` file to log "Hello World" to the console.

# --hints--

`"Hello World"` should be in the console

```js
  $.get(code + '/_api/hello-console').then(
    (data) => {
      assert.isTrue(data.passed, '"Hello World" is not in the server console');
    },
    (xhr) => {
      throw new Error(xhr.responseText);
    }
  );
```

