---
id: 587d7fb0367417b2b2512bed
title: 認識 Node 主控台
challengeType: 2
forumTopicId: 301515
dashedName: meet-the-node-console
---

# --description--

處理這些挑戰將涉及你使用以下其中一種方法來撰寫你的程式碼：

- 複製<a href="https://github.com/freeCodeCamp/boilerplate-express/" target="_blank" rel="noopener noreferrer nofollow">這個 GitHub 倉庫</a>，並在區域端完成這些挑戰。
- 使用你選擇的網站建置工具來完成專案。務必將我們 GitHub 倉庫中的所有檔案納入其中。

在開發過程中，能夠檢查你的程式碼中發生了什麼是很重要的。

Node 只是 JavaScript 環境。像客戶端 JavaScript 一樣，你可以使用主控台來顯示有用的偵錯資訊。在你的區域機器上，你會在終端機中看到主控台輸出。

我們建議在進行這些挑戰時保持終端機開啟。透過閱讀終端機中的輸出，你可以看到可能發生的任何錯誤。

在對伺服端的檔案進行更改後，必須重新啟動伺服端。

你可以在終端機使用 `Ctrl + C` 停止伺服器，並直接使用 Node（`node mainEntryFile.js`）或在 `package.json` 檔案中使用 `npm run` 的執行腳本來啟動它。

例如，`"start": "node server.js"` 腳本會使用 `npm run start` 從終端機執行。

要在檔案存檔時實作伺服端自動重啟，Node 提供了你可以添加到啟動腳本的 `--watch` 旗標 `"start": "node --watch server.js"`，或者你可以安裝像 `nodemon` 這樣的 npm 軟體包。我們將這部分留給你作為練習。

# --instructions--

修改 `myApp.js` 檔案以將 "Hello World" 記錄到主控台。

# --hints--

"Hello World" 應該出現在主控台

```js
  const response = await fetch(code + '/_api/hello-console');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(data.passed, '"Hello World" is not in the server console');
```

