---
id: 587d7fb1367417b2b2512bf1
title: 在特定路由上提供 JSON
challengeType: 2
forumTopicId: 301517
dashedName: serve-json-on-a-specific-route
---

# --description--

當 HTML 伺服器提供 HTML 時，API 則提供資料。<dfn>REST</dfn>（表述性狀態轉移）API 允許以簡單的方式交換資料，而不需要客戶端知道任何關於伺服器的細節。客戶端只需要知道資源的位置（URL）以及它想對該資源執行的動作（動詞）。當你擷取某些資訊而不修改任何內容時，會使用 GET 動詞。現今，在網路上傳輸資訊時，偏好的資料格式是 JSON。簡單來說，JSON 是一種方便將 JavaScript 物件表述為字串的方式，因此可以輕鬆傳輸。

讓我們透過創建一個在路徑 `/json` 回應 JSON 的路由來創建一個簡單的 API。你可以像平常一樣，使用 `app.get()` 方法。在路由處理常式中，使用 `res.json()` 方法，並將一個物件作為引數傳入。此方法會關閉請求-回應迴圈，並傳回資料。在背後，它會將有效的 JavaScript 物件轉換成字串，然後設定適當的標頭以告訴你的瀏覽器你正在提供 JSON，並將資料傳回。有效的物件具有通常的結構 `{key: data}`。`data` 可以是數字、字串、巢狀物件或陣列。`data` 也可以是變數或函式呼叫的結果，在這種情況下，它會在轉換成字串之前被評估。

# --instructions--

以 JSON 格式將物件 `{"message": "Hello json"}` 作為回應，回應對 `/json` 路由的 GET 請求。然後將你的瀏覽器指向 `your-app-url/json`，你應該會在螢幕上看到該訊息。

# --hints--

端點 `/json` 應該回傳 JSON 物件 `{"message": "Hello json"}`

```js
  const response = await fetch(code + '/json');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.message,
    'Hello json',
    "The '/json' endpoint does not serve the right data"
  );
```

