---
id: 587d7fb2367417b2b2512bf5
title: 從客戶端取得路由參數輸入
challengeType: 2
forumTopicId: 301513
dashedName: get-route-parameter-input-from-the-client
---

# --description--

在建置 API 時，我們必須允許使用者告訴我們他們想從我們的服務中取得什麼。 例如，如果客戶端正在請求儲存在資料庫中的使用者資訊，他們需要一種方式讓我們知道他們感興趣的是哪個使用者。 達成此結果的一種可能方法是使用路由參數。路由參數是 URL 中以斜線（/）分隔的命名區段。每個區段會擷取與其位置相符的 URL 零件的值。擷取的值可以在 `req.params` 物件中找到。

<blockquote>route_path: '/user/:userId/book/:bookId'<br>actual_request_URL: '/user/546/book/6754' <br>req.params: {userId: '546', bookId: '6754'}</blockquote>

# --instructions--

建置一個回聲伺服器，掛載於路由 `GET /:word/echo`。回應一個 JSON 物件，採用結構 `{echo: word}`。你可以在 `req.params.word` 找到要重複的字。你可以從瀏覽器的位址列測試你的路由，造訪一些符合條件的路由，例如 `your-app-rootpath/freecodecamp/echo`。

# --hints--

測試 1 ：你的 echo 伺服器應該正確重複字。

```js
  const response = await fetch(code + '/eChOtEsT/echo');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.echo,
    'eChOtEsT',
    'Test 1: the echo server is not working as expected'
  );
```

測試 2 ：你的 echo 伺服器應該正確重複字。

```js
  const response = await fetch(code + '/ech0-t3st/echo');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.echo,
    'ech0-t3st',
    'Test 2: the echo server is not working as expected'
  );
```

