---
id: 690ddb5e6ee94eb73ed172b0
title: معالجة الأخطاء في نموذج التبرع
challengeType: 25
dashedName: debug-donation-form
# Purposefully removed the demo type because this is a debugging project.
---

# --description--

أنشأت جمعية خيرية محلية موقعًا إلكترونيًا لنموذج التبرع، لكنه يحتوي على عدة مشكلات تحتاج إلى إصلاح. النموذج غير قابل للوصول ويحتوي على بعض أخطاء الصيغة النحوية في HTML.

مهمتك هي إصلاح جميع الأخطاء ليعمل النموذج بشكل صحيح ويكون متاحًا لجميع المستخدمين. أكمل العناصر في قصص المستخدم أدناه ثم انقر على "تحقق من شفرتك" لترى ما إذا كنت قد أصلحت كل الأخطاء.

**قصص المستخدم:**

1. عناصر `input` هي عناصر فارغة ولا يجب أن تحتوي على وسم إغلاق. احذف جميع وسم الإغلاق `</input>` من النموذج.
2. أضف عناصر `label` لكل حقل إدخال في النموذج ليعرف المستخدمون وظيفة كل حقل. يجب أن يتطابق نص المسمى مع النص الموجود حاليًا بجانب كل مُدخَل.
3. يجب أن يكون نوع الإدخال `Email Address:` هو `email` بدلاً من `text`.
4. اربط كل عنصر `label` بعنصر `input` المقابل له باستخدام الخاصية `for` في `label` و`id` المطابق في `input`.
5. أضف الخاصية `required` إلى حقول الإدخال النصية والبريد الإلكتروني والرقمية (لكن ليس إلى مربع الاختيار أو زر الإرسال) لضمان ملء المستخدمين للمعلومات المطلوبة.

# --hints--

يجب ألا يحتوي شفرتك على أي وسم إغلاق `</input>`.

```js
assert.notMatch(code, /<\/input>/);
```

يجب أن يحتوي نموذجك على خمسة عناصر `input` بالضبط.

```js
assert.lengthOf(document.querySelectorAll("input"), 5);
```

يجب أن يحتوي نموذجك على أربعة عناصر `label` بالضبط.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

يجب أن يحتوي أول `label` على النص `Full Name:`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[0]?.textContent, /Full\s*Name\s*:/i);
```

يجب أن يحتوي أول `input` على الخاصية `required`.

```js
const firstInput = document.querySelector("input[name='name']");
assert.isTrue(firstInput?.hasAttribute("required"));
```

يجب أن يحتوي ثاني `label` على النص `Email Address:`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[1]?.textContent, /Email\s*Address\s*:/i);
```

يجب أن يحتوي ثاني `input` على الخاصية `required`.

```js
const secondInput = document.querySelector("input[name='email']");
assert.isTrue(secondInput?.hasAttribute("required"));
```

يجب أن يحتوي ثالث `label` على النص `Donation Amount ($):`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[2]?.textContent, /Donation\s*Amount\s*\(\$\)\s*:/i);
```

يجب أن يحتوي ثالث `input` على الخاصية `required`.

```js
const thirdInput = document.querySelector("input[name='amount']");
assert.isTrue(thirdInput?.hasAttribute("required"));
```

يجب أن يحتوي رابع `label` على النص `Subscribe`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[3]?.textContent, /Subscribe/i);
```

يجب ربط كل عنصر `input` بعنصر `label` باستخدام الخاصيتين `for` و`id`.

```js 
const els = document.querySelectorAll('input:not([type="submit"])'); 
assert.isNotEmpty(els); 
els.forEach(input => { 
  const label = document.querySelector(`label[for="${input.id}"]`); 
  assert.exists(label);  
  assert.equal(label.tagName, "LABEL"); 
})
``` 

يجب أن يكون نوع `input` الخاص بعنوان البريد الإلكتروني هو `email`.

```js
const emailInput = document.querySelector("input[name='email']");
assert.strictEqual(emailInput?.getAttribute("type"), "email");
```

يجب ألا يحتوي مربع اختيار `input` على الخاصية `required`.

```js
const checkboxInput = document.querySelector("input[name='newsletter']");
assert.isFalse(checkboxInput?.hasAttribute("required"));
```

يجب ألا يحتوي زر الإرسال `button` على الخاصية `required`.

```js
const submitButton = document.querySelector("input[type='submit']");
assert.isFalse(submitButton?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    
    Full Name:
    <input type="text" name="name"></input>

    Email Address:
    <input type="text" name="email">

    Donation Amount ($):
    <input type="number" name="amount"></input>

    <input type="checkbox" name="newsletter"></input>
    Subscribe

    <input type="submit" value="Send"></input>
  </form>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    <label for="full-name">Full Name:</label>
    <input id="full-name" type="text" name="name" required>

    <label for="email">Email Address:</label>
    <input id="email" type="email" name="email" required>

    <label for="amount">Donation Amount ($):</label>
    <input id="amount" type="number" name="amount" required>

    <label for="newsletter">Subscribe</label>
    <input id="newsletter" type="checkbox" name="newsletter">

    <input type="submit" value="Send">
  </form>
</body>
</html>
```
