---
id: 69a016ac6edea7d8c8020c98
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

Ora crea un loop `while` con `current.next !== null` come condizione. Questo significa che il loop deve continuare a eseguire finché il nodo corrente ha un altro nodo dopo di sé.

# --hints--

Dovresti creare un loop `while` con la condizione `current.next !== null`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /while\s+\(current\.next\s*!==\s*null\)\s*\{\s*/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        --fcc-editable-region--

        --fcc-editable-region--
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
