---
id: 69a00cd4ec513345d40e95ec
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

La prossima cosa su cui lavorerai è una funzione che ti permette di verificare se la lista è vuota.

Per farlo, crea una funzione chiamata `isEmpty` che prende un parametro `list`. All’interno della funzione, restituisci un controllo per verificare se la proprietà `length` di `list` è strettamente uguale a `0`.

# --hints--

Dovresti creare una funzione `isEmpty`.

```js
assert.isFunction(isEmpty)
```

La tua funzione `isEmpty` dovrebbe avere un parametro `list`.

```js
const isEmptyParams = __helpers.getFunctionParams(isEmpty.toString());
assert.deepEqual(isEmptyParams[0].name, "list");
```

La tua funzione `isEmpty` dovrebbe restituire `list.length === 0`.

```js
const emptyList = { head: null, length: 0 };
const nonEmptyList = { head: {}, length: 1 };

assert.isTrue(isEmpty(emptyList));
assert.isFalse(isEmpty(nonEmptyList));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
