---
id: 69a00eaa43458e618331873c
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Ora lavorerai sulla funzione che aggiungerà un nodo alla lista.

Crea una funzione `add` con i parametri `list` e `element`.

# --hints--

Dovresti creare una funzione `add`.

```js
assert.isFunction(add)
```

La tua funzione `add` dovrebbe avere i parametri `list` e `element`.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.add.parameters;
assert.equal(parameters[0].toString(), "list");
assert.equal(parameters[1].toString(), "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
