---
id: 69a045fd236d0d973a59f103
title: Passaggio 19
challengeType: 1
dashedName: step-19
---

# --description--

Se `previous` è `null`, significa che l'elemento da rimuovere si trova all'inizio della lista, quindi dovresti aggiornare il `head` per puntare al nodo successivo per rimuovere il primo elemento.

Per farlo, aggiungi una clausola `else`. Al suo interno, imposta `list.head` su `current.next`.

# --hints--

Dovresti aggiungere una clausola `else` alla tua istruzione `if`.

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*/);
```

Dovresti impostare `list.head` su `current.next` all'interno della clausola `else`.

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*list\.head\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }

    if(previous !== null) {
        previous.next = current.next;
    --fcc-editable-region--
    } 
        
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
