---
id: 69a04088e600bd8074a9852f
title: Passaggio 17
challengeType: 1
dashedName: step-17
---

# --description--

Se `current` è `null`, significa che è la fine della lista, quindi la funzione dovrebbe smettere di eseguire.

Crea un'istruzione `if` che verifichi se `current` è strettamente uguale a `null`. All'interno dell'istruzione `if`, usa `return` per uscire dalla funzione.

# --hints--

Dovresti creare un'istruzione `if` con la condizione `current === null`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(current\s*===\s*null\s*\)\s*\{\s*/);
```

Dovresti usare la parola chiave `return` all'interno della tua istruzione `if`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(current\s*===\s*null\s*\)\s*\{\s*return;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
