---
id: 69a012c6224fac85832247dd
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Hai notato i risultati `true`, `{ head: null, length: 0 }` e `true` nella console? Succede perché non hai assegnato nulla alla lista `myList` e non hai incrementato la proprietà `length`.

Vai avanti e assegna `node` a `list.head`, e sotto di quello, incrementa la proprietà `length` di `list`.

# --hints--

Dovresti assegnare `node` a `list.head`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /list\.head\s*=\s*node;?/);
```

Dovresti incrementare la proprietà `length` di `list`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /list\.length\+\+;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    --fcc-editable-region--

    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
