---
id: 69a014bd9b1730e94136279f
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Se la lista non è vuota, allora devi trovare la sua fine, aggiungere il nuovo elemento lì e farlo puntare a `null`. Questo significa che devi iterare la lista a partire da `head`.

Per iniziare, apri una clausola `else`. Al suo interno, usa `let` per creare una variabile `current` e impostarla su `list.head`.

# --hints--

Dovresti avere una clausola `else`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{/);
```

Dovresti creare una variabile `current` impostata su `list.head` all’interno della clausola `else`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{\s*let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    --fcc-editable-region--
    } 
    --fcc-editable-region--

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
