---
id: 69a0102e761a245bb1cbd8e7
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

A questo punto, dovresti testare la lista concatenata per capirla meglio.

Crea una variabile `myList` e impostala su `initList()`. Sotto di essa, registra una chiamata di `isEmpty()` usando `myList` come parametro nella console. Poi, subito dopo, chiama `add` con `myList` e un numero a tua scelta, quindi registra `myList` nella console. Sotto, registra un'altra chiamata di `isEmpty()` nella console usando di nuovo `myList` come parametro.

# --hints--

Dovresti creare una variabile `myList` impostata su `initList()`.

```js
assert.deepEqual(myList, initList());
```

Dovresti registrare una chiamata a `isEmpty(myList)` nella console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(isEmpty\(myList\)\);?/)
```

Dovresti chiamare `add` con `myList` e un numero a tua scelta come parametri.

```js
assert.match(__helpers.removeJSComments(code), /add\s*\(\s*myList\s*,\s*\d+\s*\)/);
```

Dovresti registrare la tua variabile `myList` nella console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?/)
```

Dovresti registrare di nuovo `isEmpty(myList)` nella console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?\s*console\.log\(isEmpty\(myList\)\);?/)
```


# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
