---
id: 69a00a051b9a18694d00c16b
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop implementerai una lista concatenata.

Per cominciare, crea una funzione chiamata `initList`. Al suo interno, restituisci un oggetto con una proprietà `head` impostata su `null` e una proprietà `length` di `0`.

Questa è la funzione che chiamerai per inizializzare una lista concatenata. La `head` rappresenta il primo elemento della lista, inizialmente è `null` perché non ci sono elementi nella lista. La `length` tiene traccia del numero di nodi (elementi) nella lista, che incrementerai o decrementerai quando aggiungi o rimuovi un nodo dalla lista.

# --hints--

Dovresti creare una funzione `initList`.

```js
assert.isFunction(initList)
```

La tua funzione `initList` dovrebbe restituire un oggetto.

```js
assert.isObject(initList());
```

La tua funzione `initList` dovrebbe restituire un oggetto con una `head` impostata su `null`.

```js
assert.propertyVal(initList(), "head", null);
```

La tua funzione `initList` dovrebbe restituire un oggetto con una `length` impostata su `0`.

```js
assert.propertyVal(initList(), "length", 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
