---
id: 69a03aaa1e6fe8152788ce6a
title: Passaggio 14
challengeType: 1
dashedName: step-14
---

# --description--

All'interno della funzione `remove`, usa `let` per creare una variabile `previous` impostata su `null`, e di nuovo `let` per creare una variabile `current` impostata su `list.head`. Queste sono le variabili che userai nel loop che attraverserà la lista.

`current` tiene traccia del nodo attualmente controllato nella lista, partendo dalla testa e avanzando di un nodo alla volta. `previous`, invece, tiene traccia del nodo subito prima di quello corrente.

Questo è importante quando si rimuove un nodo perché devi aggiornare il valore next del nodo precedente in modo che punti al nodo successivo a quello rimosso, per mantenere la lista collegata.

# --hints--

Dovresti usare `let` per creare una variabile `previous` impostata su `null` all'interno della tua funzione `remove`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+previous\s*=\s*null;?/);
```

Dovresti usare `let` per creare una variabile `current` impostata su `list.head` all'interno della tua funzione `remove`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
