---
id: 69a013a86c97dac3784b2636
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Ora, il log dovrebbe mostrare qualcosa come `{ head: { element: 42, next: null }, length: 1 }` nella console.

Tuttavia, questo non è il modo corretto di gestire il caso in cui il nodo aggiunto è il primo elemento della lista, cioè la testa.

Per farlo correttamente, usa un'istruzione if per verificare se la lista è vuota chiamando la funzione `isEmpty` con `list`. Se è vero, allora sposta la tua riga `list.head = node` all'interno del corpo dell'istruzione `if`.

# --hints--

Dovresti avere un'istruzione `if` con la condizione `isEmpty(list)`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'if(isEmpty(list)){');
```

Dovresti spostare la tua riga `list.head = node` all'interno dell'istruzione `if`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'if(isEmpty(list)){list.head=node');
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

--fcc-editable-region--
    list.head = node;
--fcc-editable-region--
    list.length++;
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
console.log(myList);
console.log(isEmpty(myList));
```
