---
id: 69a01864d343dc2f2484d131
title: Passaggio 11
challengeType: 1
dashedName: step-11
---

# --description--

Dopo aver trovato l’ultimo nodo, devi allegare il nuovo nodo a quello. Per farlo, imposta `current.next` su `node` all’interno del loop (iterazione) `else`, ma ancora dentro la clausola `else`.

# --hints--

Dovresti impostare `current.next` su `node` all’interno del tuo loop (iterazione) `while`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\.next\s*=\s*node;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        --fcc-editable-region--

        --fcc-editable-region--
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
