---
id: 69a0179a85b00b6127e0af45
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

All’interno del loop `while`, aggiorna `current` a `current.next` in modo che il loop iteri la lista un nodo alla volta e si fermi quando raggiunge il nodo con una proprietà `next` di valore `null`, cioè la fine della lista.

# --hints--

Dovresti aggiornare il valore della tua variabile `current` a `current.next` all’interno del loop `while`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            --fcc-editable-region--

            --fcc-editable-region--
        }
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
