---
id: 69a00f673f171e29159e3ad8
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Questo nodo dovrebbe avere due proprietà, l'elemento da aggiungere e un puntatore al nodo successivo.

All'interno della funzione add, crea un oggetto chiamato `node`, che deve avere una proprietà chiamata `element` e una proprietà `next` impostata su `null`, che è il puntatore all'elemento successivo della lista.

# --hints--

Dovresti creare un oggetto chiamato `node` all'interno della tua funzione `add`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /const|var|let\s+node\s*=\s*{/);
```

Il tuo oggetto `node` dovrebbe avere una proprietà chiamata `element` impostata su `element`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /element\s*:\s*element/);
```

Il tuo oggetto `node` dovrebbe avere una proprietà chiamata `next` impostata su `null`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /next\s*:\s*null/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}
```
