---
id: 69a04466917c66b132323a1b
title: Passaggio 18
challengeType: 1
dashedName: step-18
---

# --description--

Se `previous` non è `null`, significa che l'elemento da rimuovere non si trova all'inizio ma da qualche parte nella lista, quindi il nodo corrente deve essere saltato collegando direttamente il nodo precedente a quello successivo.

Per farlo, crea un'altra istruzione `if` che verifichi se previous è strettamente diverso da `null`. Se è così, imposta `previous.next` su `current.next`.

# --hints--

Dovresti creare un'altra istruzione `if` con la condizione `previous !== null`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*/);
```

Dovresti impostare `previous.next` su `current.next`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*previous\.next\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }
--fcc-editable-region--
    
--fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
