---
id: 642db8c409d9991d0b3b2f0d
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dein Projekt beginnt mit einem einfachen HTML-Container und einigem entsprechenden CSS. Deine erste Aufgabe besteht darin, die Zellen für dein Spreadsheet zu programmieren.

Das globale `window`-Objekt stellt das Browserfenster (oder die Registerkarte) dar. Es hat eine `onload`-Eigenschaft, mit der du das Verhalten festlegen kannst, wenn das Fenster die gesamte Seite, einschließlich Stylesheets und Skripts geladen hat.

Beginne, indem du die `onload`-Eigenschaft von `window` auf eine Arrow-Funktion ohne Parameter setzt. In the function, declare a `container` variable and assign it the value of getting the element by the `id` of `"container"`.

# --hints--

Du solltest auf die `onload`-Eigenschaft des `window`-Objekts zugreifen.

```js
assert.match(code, /window\.onload/);
```

Du solltest die `onload`-Eigenschaft auf eine Funktion setzen.

```js
assert.isFunction(window.onload);
```

Du solltest Arrow-Syntax verwenden.

```js
assert.match(code, /window\.onload\s*=\s*\(/);
```

Deine `onload`-Funktion sollte keinen Parameter verwenden.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)/);
```

Du solltest eine `container`-Variable in deiner `onload`-Funktion deklarieren.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*(?:let|var|const)\s+container/);
```

Deine `container`-Variable sollte mit `const` deklariert werden.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container/);
```

Du solltest `document.getElementById()` verwenden

```js
assert.match(code, /document\.getElementById\(/);
```

You should get the element with the `id` of `"container"`.

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

Du solltest dem `container` das `#container`-Element zuweisen.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
