---
id: 66be24cb4144f955b6bcc550
title: Construa uma página de jogos de cartas
challengeType: 25
dashedName: build-a-page-of-playing-cards
demoType: onClick
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma página da web que exiba pelo menos três cartas de baralho.
1. Você deve ter um elemento `main` com um ID de `playing-cards`.
1. Dentro do elemento `#playing-cards`, você deve ter pelo menos três elementos `div`, cada um com uma classe de `card`.
1. Dentro de cada elemento `.card`, você deve ter três elementos `div`, o primeiro com uma classe de `left`, o segundo com uma classe de `middle` e o terceiro com uma classe de `right`.
1. O elemento `#playing-cards` deve usar flexbox para centralizar horizontalmente os elementos filhos, permitir que eles se agrupem e colocar um espaço de `20px` entre eles.
1. Cada um dos elementos `.card` deve usar flexbox para justificar os elementos filhos usando `space-between`.
1. Cada um dos elementos `.left` deve usar propriedades de item do flexbox para se alinhar no início do contêiner pai.
1. Cada um dos elementos `.middle` deve usar propriedades de item do flexbox para se alinhar ao centro do contêiner pai.
1. Cada um dos elementos `.right` deve usar propriedades de item do flexbox para se alinhar ao final do contêiner pai.
1. Cada um dos elementos `.middle` deve usar o flexbox para exibir os elementos filhos em uma coluna.

Aqui estão alguns caracteres que você pode copiar e colar para usar em seus cartões, se desejar: `♠`, `♣`, `♥` e `♦`.

# --hints--

Você deve ter um elemento `main` com um `id` de `playing-cards`.

```js
assert.exists(document.querySelector('main#playing-cards'));
```

Você deve ter pelo menos três elementos `div` com uma classe de `card` dentro do elemento `#playing-cards`.

```js
assert.isAtLeast(document.querySelectorAll('#playing-cards > div.card')?.length, 3);
```

Cada um dos elementos `.card` deve ter uma largura e altura.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  assert.isAbove(el.clientWidth, 0);
  assert.isAbove(el.clientHeight, 0);
});
```

Cada um dos elementos `.card` deve ter exatamente três elementos `div` como filhos.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  const children = el.children;
  assert.lengthOf(children, 3);
  Array.from(children).forEach(child => {
    assert.equal(child.tagName, 'DIV');
  });
});
```

Você deve ter uma `div` com uma classe de `left` dentro de cada elemento `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.left')));
```

Você deve ter uma `div` com uma classe de `middle` dentro de cada elemento `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.middle')));
```

Você deve ter uma `div` com uma classe de `right` dentro de cada elemento `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.right')));
```

Você deve ter um seletor `#playing-cards` que defina o `display` dos elementos como `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.display, 'flex');
```

Você deve ter um seletor `#playing-cards` que defina o `justify-content` dos elementos como `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.justifyContent, 'center');
```

Você deve ter um seletor `#playing-cards` que defina o `flex-wrap` dos elementos como `wrap`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.flexWrap, 'wrap');
```

Você deve ter um seletor `#playing-cards` que defina o `gap` dos elementos como `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.gap, '20px');
```

Você deve ter um seletor `.card` que defina o `display` dos elementos como `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.display, 'flex');
```

Você deve ter um seletor `.card` que defina o `justify-content` dos elementos como `space-between`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.justifyContent, 'space-between');
```

Você deve ter um seletor `.left` que defina o `align-self` dos elementos como `flex-start`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left')?.alignSelf, 'flex-start');
```

Você deve ter um seletor `.middle` que defina o `align-self` dos elementos como `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.alignSelf, 'center');
```

Você deve ter um seletor `.right` que defina o `align-self` dos elementos como `flex-end`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.right')?.alignSelf, 'flex-end');
```

Você deve ter um seletor `.middle` que define o `display` de seus elementos como `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.display, 'flex');
```

Você deve ter um `flex-direction` definido como `column` na sua classe `.middle`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Playing Cards</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Straight</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="playing-cards">
    <div class="card">
      <div class="left">
        <div>A</div>
        <div>♠</div>
      </div>
      <div class="middle">
        <div>♠</div>
      </div>
      <div class="right">
        <div class="flipped">♠</div>
        <div class="flipped">A</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>2</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div>♣</div>
        <div class="flipped">♣</div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">2</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>3</div>
        <div>♥</div>
      </div>
      <div class="middle">
        <div>♥</div>
        <div>♥</div>
        <div class="flipped">♥</div>
      </div>
      <div class="right">
        <div class="flipped">♥</div>
        <div class="flipped">3</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>4</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♣</div>
          <div>♣</div>
        </div>
        <div class="middle-section">
          <div class="flipped">♣</div>
          <div class="flipped">♣</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">4</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>5</div>
        <div>♦</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♦</div>
        <div class="flipped">5</div>
      </div>
    </div>
  </main>
</body>

</html>
```

```css
body {
  background-color: Thistle;
  min-height: 100vh;
}

.red {
  color: red;
}

.flipped {
  transform: rotate(180deg);
}

#playing-cards {
  margin-top: 100px;
  display: flex;
  justify-content: center;
  gap: 20px;
  flex-wrap: wrap;
}

.card {
  background-color: white;
  border: 1px solid black;
  border-radius: 10px;
  box-shadow: 2px 2px 2px gray;
  min-width: 250px;
  max-width: 250px;
  height: 350px;
  font-size: 30px;
  font-weight: bold;
  padding: 5px;
  box-sizing: border-box;
  display: flex;
  justify-content: space-between;
  text-align: center;
}

.left {
  align-self: flex-start;
}

.middle {
  align-self: center;
  display: flex;
  font-size: 80px;
  flex-direction: column;
}

.right {
  align-self: flex-end;
}

.middle-section {
  display: flex;
  justify-content: center;
}
```
