---
id: 5d822fd413a79914d39e98c9
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Welcome to the CSS Variables Skyline project! Start by adding the `!DOCTYPE html` declaration at the top of the document so the browser knows what type of document it's reading.

# --hints--

代碼應包含 `DOCTYPE` 聲明。

```js
assert(code.match(/<!DOCTYPE/gi));
```

應該在 `DOCTYPE` 聲明之後包含一個空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

應該將文檔類型定義爲 `html`。

```js
assert(code.match(/html/gi));
```

你應該在類型之後使用 `>` 關閉 `DOCTYPE` 聲明。

```js
assert(code.match(/html\s*>/gi));
```

`DOCTYPE` 聲明應位於 HTML 的開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
