---
id: 6823e036ea4b71553558c01b
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Observe que cada título parece um pouco diferente na pré-visualização. Mesmo com o parágrafo.

Em seguida, adicione outro elemento `p` que exiba o seguinte texto:

```md
HTML represents the content and structure of a webpage
```

# --hints--

Você deve ter um segundo elemento `p`.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

Seu elemento `p` deve conter o texto `HTML represents the content and structure of a webpage`. Você pode copiar o texto das instruções.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*HTML\s*represents\s*the\s*content\s*and\s*structure\s*of\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
--fcc-editable-region--

--fcc-editable-region--
```
