---
id: 6823e637c1c0ed56f781b4fc
title: Passo 11
challengeType: 0
dashedName: step-11
---

# --description--

Finalmente, JavaScript torna sua página web interativa — ele permite que você diga à página o que fazer quando alguém clica em um botão, envia um formulário ou muitas outras coisas.

Para a última etapa do workshop, adicione outro elemento `h3` e `p` à página descrevendo JavaScript.

Primeiro, adicione um elemento `h3` com o texto:

```md
Introduction to JavaScript
```

Então, abaixo daquele elemento `h3`, adicione um elemento `p` com o seguinte texto:

```md
JavaScript adds interactivity to a webpage
```

# --hints--

Você deve adicionar um terceiro elemento `h3` à página.

```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

O texto do seu elemento `h3` deve ser `Introduction to JavaScript`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*JavaScript\s*\<\/h3\s*\>/i);
```

Você deve adicionar um quarto elemento `p` à página.

```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

Seu elemento `p` deve conter o texto `JavaScript adds interactivity to a webpage`

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*JavaScript\s*adds\s*interactivity\s*to\s*a\s*web\s*page\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

<h3>Introduction to JavaScript</h3>
<p>JavaScript adds interactivity to a webpage</p>
```
