---
id: 6823d9ac8bdc3853df65a1ff
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Existem seis elementos de título em HTML: `h1` a `h6`.Eles são usados para mostrar a importância das seções na sua página, com `h1` sendo o mais importante e `h6` o menos importante.

Abaixo do seu elemento `p`, adicione um título `h3` com o texto:

```md
Introduction to HTML
```

# --hints--

Seu elemento `h3` deve ter uma tag de abertura `<h3>`.

```js
assert.exists(document.querySelector("h3"));
```

Seu elemento `h3` deve ter uma tag de fechamento `</h3>`.

```js
assert.match(code, /<\/h3\s*\>/);
```

O texto do seu elemento `h3` deve ser `Introduction to HTML`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

--fcc-editable-region--

--fcc-editable-region--
```
