---
id: 6823d6244511f252c8300eed
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Quando você precisar adicionar um parágrafo a uma página da web, pode usar o elemento `p` assim:

```html
<p>This is a paragraph element.</p>
```

Transforme `Learn the skills to become a full stack developer` em um elemento de parágrafo.

# --hints--

O elemento `p` deve ter uma tag de abertura.As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert.exists(document.querySelector("p"));
```

Seu elemento `p` deve ter uma tag de fechamento. Tags de fechamento têm esta sintaxe: `</elementName>`.

```js
assert.match(code, /<\/p\s*\>/);
```

O texto do seu elemento `p` deve ser `Learn the skills to become a full stack developer`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*Learn\s*the\s*skills\s*to\s*become\s*a\s*full\s*stack\s*developer\s*\<\/p\s*\>/i);
```

Seu elemento `p` deve estar abaixo do seu elemento `h2`. 

```js
assert.exists(document.querySelector("h2 + p"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
--fcc-editable-region--
Learn the skills to become a full stack developer
--fcc-editable-region--
```
