---
id: 6823ac607bfdbc46331b2559
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML significa HyperText Markup Language. É o código que define a estrutura e o conteúdo de uma página da web. Este é o seu editor de código, onde você escreverá HTML.

Encontre a linha 1 no editor e digite este texto:

```md
Welcome to freeCodeCamp
```

Quando terminar, clique no botão "check your code" para ver se está correto.

# --hints--

Você deve ter o texto `Welcome to freeCodeCamp` no seu editor. Verifique a ortografia.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /welcome\s*to\s*freecodecamp/i)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
