---
id: 682ba2318000b62f179bdf04
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

HTML é composto por elementos. O primeiro que você usará é o elemento `h1`:

```html
<h1>Welcome to freeCodeCamp</h1>
```

Começa com uma tag de abertura (`<h1>`), termina com uma tag de fechamento (`</h1>`), e tem o texto que será exibido entre as tags.

Transforme seu texto `Welcome to freeCodeCamp` em um elemento `h1` adicionando uma tag de abertura na frente dele e uma tag de fechamento depois dele.

# --hints--

Seu elemento `h1` deve ter uma tag de abertura `<h1>`.

```js
assert.exists(document.querySelector("h1"));
```

Seu elemento `h1` deve ter uma tag de fechamento `</h1>`.

```js
assert.match(code, /<\/h1\s*\>/);
```

Seu elemento `h1` deve ser assim: `<h1>Welcome to freeCodeCamp</h1>`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h1\s*\>\s*Welcome\s*to\s*freeCodeCamp\s*\<\/h1\s*\>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
Welcome to freeCodeCamp
--fcc-editable-region--
```
