---
id: 682cd206883fc7b25eb539c5
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Observe que o HTML que você escreve no editor aparece na pré-visualização.Neste workshop, você escreverá o HTML para uma página parcial de currículo.

Abaixo do seu elemento `h1`, digite o seguinte na linha vazia:

```md
Full Stack Curriculum
```

# --hints--

Você deve ter o texto `Full Stack Curriculum` no seu editor. Verifique a ortografia.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /full\s*stack\s*curriculum/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
--fcc-editable-region--

--fcc-editable-region--
```
