---
id: 683921c4769fd23dbadec2fe
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Enquanto o HTML define a estrutura e o conteúdo de uma página web, o CSS é usado para adicionar estilo — coisas como cores, fontes, espaçamento e layout.

Abaixo do `h3` que você acabou de adicionar, adicione outro elemento de parágrafo com o texto:

```md
CSS is used to style a webpage
```

# --hints--

Você deve adicionar um segundo elemento `h3` à página.

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

O texto do seu elemento `h3` deve ser `Introduction to CSS`. 

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

Você deve adicionar um terceiro elemento `p` à página.

```js
assert.lengthOf(document.querySelectorAll("p"), 3);
```

Seu elemento `p` deve conter o texto `CSS is used to style a webpage`. Você pode copiar o texto das instruções.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*CSS\s*is\s*used\s*to\s*style\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

--fcc-editable-region--
<h3>Introduction to CSS</h3>

--fcc-editable-region--
```
