---
id: 614ccc21ea91ef1736b9b578
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Te damos la bienvenida a la primera parte del Quiz de Accesibilidad. Como te estás convirtiendo en un desarrollador experimentado de HTML y CSS, te hemos proporcionado una plantilla básica inicial.

Comienza este viaje de accesibilidad proporcionando un atributo `lang` a tu elemento `html`. Esto ayudará a los lectores de pantalla a identificar el idioma de la página.

# --hints--

Debes añadir un atributo `lang` al elemento `html`. _Consejo: Puedes usar el valor `en` para inglés._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
