---
id: 686386e76197458b7dfd52ee
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

All'interno del `ul`, crea due elementi di lista. Nel primo elemento di lista, crea un elemento `span` e aggiungi l'entità HTML `&#10003;` come suo testo. Subito dopo quel `span` ma all'interno dell'elemento di lista, aggiungi il testo `Ad-supported streaming`.

Nel secondo elemento di lista, crea un elemento `span` e aggiungi l'entità `&#10003;` come suo testo. Subito dopo quel `span`, ma all'interno dell'elemento di lista, aggiungi il testo `Curated playlists`.

# --hints--

Dovresti avere due elementi di lista all'interno del tuo elemento `ul`.

```js
assert.lengthOf(document.querySelectorAll("ul > li"), 2)
```

Il tuo primo elemento `li` dovrebbe contenere un elemento `span` con l'entità carattere HTML `&#10003;` al suo interno.

```js
const listItems = document.querySelectorAll("ul > li")
const firstSpan = listItems[0].querySelector("span")

assert.equal(firstSpan?.textContent, "✓")
```

Il tuo primo elemento `li` dovrebbe avere il testo `Ad-supported streaming` subito dopo l'elemento `span`.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[0].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Ad-supported streaming");
```

Il tuo secondo elemento `li` dovrebbe contenere un elemento `span` con l'entità carattere HTML `&#10003;` al suo interno.

```js
const listItems = document.querySelectorAll("ul > li")
const secondSpan = listItems[1].querySelector("span")

assert.equal(secondSpan?.textContent, "✓")
```

Il tuo secondo elemento `li` dovrebbe avere il testo `Curated playlists` subito dopo l'elemento `span`.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[1].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Curated playlists");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>Start exploring millions of songs with basic features and ads.</p>
            <ul>
            --fcc-editable-region--
              
            --fcc-editable-region--
            </ul>
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
