---
id: 686383abac6c625a84458ea9
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo workshop imparerai di più su Tailwind CSS usando le sue classi di utilità per la griglia per costruire un componente di prezzi.

Il codice boilerplate HTML di base è stato fornito per te. Inizia creando un elemento `h1` con il testo `Choose your listening plan`.

# --hints--

Dovresti creare un elemento `h1`.

```js
assert.exists(document.querySelector("h1"));
```

Il tuo elemento `h1` dovrebbe avere il testo `Choose your listening plan`.

```js
assert.equal(document.querySelector("h1")?.textContent, "Choose your listening plan");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </main>
  </body>
</html>
```
