---
id: 686386e76197458b7dfd52ef
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

Ora creerai un pulsante CTA. Una call to action (CTA) è un pulsante o un collegamento progettato per spingere l’utente a compiere un’azione specifica, in questo caso, iniziare ad ascoltare.

All’esterno del `div` più interno, crea un elemento àncora con un `href` impostato su `#` e il testo `Start listening`.

Con questo si conclude l’HTML per la prima scheda dei prezzi.

# --hints--

Dovresti creare un elemento àncora.

```js
assert.exists(document.querySelector("a"))
```

Il tuo elemento àncora dovrebbe avere un `href` impostato su `#`.

```js
const link = document.querySelector("a")
assert.equal(link?.getAttribute("href"), "#")
```

Il tuo elemento àncora dovrebbe avere il testo `Start listening`.

```js
assert.equal(document.querySelector("a")?.textContent, "Start listening")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          --fcc-editable-region--
          
          --fcc-editable-region-- 
        </div>
      </div>
    </main>
  </body>
</html>
```
