---
id: 686386e76197458b7dfd52eb
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Crea un elemento `div` sotto il `h1`. All’interno di quel `div`, crea un altro elemento `div` e annida un ulteriore elemento `div` al suo interno.

E dentro l’elemento `div` più interno, crea un elemento `h2` con il testo `Listener`. Questo rappresenterà il primo piano tariffario.

# --hints--

Dovresti creare un elemento `div`.

```js
assert.exists(document.querySelector("div"));
```

Dovresti creare un altro elemento `div` dentro il `div` esterno.

```js
assert.exists(document.querySelector("div > div"));
```

Dovresti creare un altro elemento `div` dentro il `div` interno.

```js
assert.exists(document.querySelector("div > div > div"));
```

Dovresti creare un `h2` con il testo `Listener` dentro l’elemento `div` più interno.

```js
assert.equal(document.querySelector("div > div > div > h2")?.textContent, "Listener");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
