---
id: 686386e76197458b7dfd52ec
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Sotto il `h2`, crea un elemento `p` con il testo `$0`. Poi, all'interno di quell'elemento `p` e dopo il testo, crea un elemento `span` con il testo `/month`.

# --hints--

Dovresti creare un elemento `p` con il testo `$0`.

```js
const p = document.querySelector("p");
const pTextBeforeSpan = p.childNodes[0].textContent.trim();

assert.equal(pTextBeforeSpan, "$0");
```

Dovresti creare un elemento `span` con il testo `/month` all'interno dell'elemento `p`.

```js
const spanText = document.querySelector("p > span")?.textContent;
assert.equal(spanText, "/month");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            --fcc-editable-region--
            
            --fcc-editable-region--
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
