---
id: 66e45c8140f9fda5c105ae26
title: Costruisci una pagina volantino per un evento
challengeType: 25
dashedName: build-an-event-flyer-page
demoType: onClick
---

# --description--

**Obiettivo:** Completa le seguenti storie utente e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un elemento `header` all'interno del body.
1. Il tuo elemento `header` dovrebbe contenere un'immagine per il tuo evento e un `h1` con il titolo del tuo evento, in quest'ordine. Puoi usare questa immagine se vuoi: `https://cdn.freecodecamp.org/curriculum/labs/event.jpg`
1. Dovresti avere un elemento `main` all'interno del body.
1. Il tuo elemento `main` dovrebbe contenere almeno due elementi `section` che mostrano le funzionalità dell'evento.
1. Ciascuno dei tuoi elementi `section` dovrebbe contenere un `h2` al suo interno.
1. Il tuo body dovrebbe avere una spaziatura interna (padding) superiore e inferiore di `50px`.
1. Il tuo body dovrebbe avere un margine superiore e inferiore di `0` e un margine sinistro e destro che lo centri.
1. Il tuo body dovrebbe avere una larghezza impostata in relazione alla larghezza del viewport.
1. Il tuo body dovrebbe usare la funzione `calc` per impostare la proprietà `min-height` al 100% dell'altezza del viewport meno tutta la spaziatura interna applicata in alto e in basso al body.
1. Dovresti avere almeno un `hr` all'interno del tuo volantino.
1. Dovresti impostare la `width` di tutti gli elementi `hr` e `section` a un valore percentuale relativo al loro genitore.

**Nota:** Assicurati di collegare il tuo foglio di stile nel tuo HTML e di applicare il CSS.

# --hints--

Dovresti avere un elemento `header` all'interno dell'elemento `body`.

```js
assert.exists(document.querySelector('body > header'));
```

Dovresti avere un elemento `img` come primo figlio in `header`.

```js
assert.equal(document.querySelector('header')?.children?.[0].tagName, 'IMG');
```

Dovresti avere un elemento `h1` come secondo figlio in `header`.

```js
assert.equal(document.querySelector('header')?.children?.[1].tagName, 'H1');
```

Dovresti avere un solo elemento `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Il tuo `h1` dovrebbe contenere il titolo del tuo evento.

```js
assert.isAtLeast(document.querySelector('h1').innerText.length, 1);
```

Dovresti avere un elemento `main` all'interno dell'elemento `body`.

```js
assert.exists(document.querySelector('body > main'));
```

Dovresti avere almeno due elementi `section` all'interno del tuo elemento `main`.

```js
assert.isAtLeast(document.querySelectorAll('main > section').length, 2);
```

Ciascuno dei tuoi elementi `section` dovrebbe contenere un `h2`.

```js
const sections = document.querySelectorAll('main > section');
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(sections.length, 1);
assert.equal(sections.length, headers.length);
```

I tuoi elementi `h2` non dovrebbero essere vuoti.

```js
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(headers.length, 1);
headers.forEach(header => assert.notEqual(header.innerHTML.trim(), ''));
```

Il tuo elemento `body` dovrebbe avere una spaziatura interna superiore e inferiore di `50px`.

```js
const body = window.getComputedStyle(document.querySelector('body'));
assert.equal(body?.paddingTop, '50px');
assert.equal(body?.paddingBottom, '50px');
```

Il tuo elemento `body` dovrebbe avere un margine superiore e inferiore di `0`.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginTop, '0px');
assert.equal(body?.marginBottom, '0px');
```

Il tuo elemento `body` dovrebbe avere un margine sinistro e destro di `auto`.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginRight, 'auto');
assert.equal(body?.marginLeft, 'auto');
```

Il tuo elemento `body` dovrebbe impostare la larghezza usando `vw`.

```js
assert.isTrue(new __helpers.CSSHelp(document).getStyle('body')?.width?.endsWith('vw'));
```

Il tuo `body` dovrebbe usare `calc` per impostare la sua `min-height` a `100vh - 100px`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('body')?.minHeight, ['calc(-100px + 100vh)', 'calc(100vh - 100px)']);
```

Dovresti avere almeno un elemento `hr`.

```js
assert.isAtLeast(document.querySelectorAll('hr')?.length, 1);
```

La larghezza dei tuoi elementi `hr` dovrebbe essere impostata usando un valore percentuale.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'hr' && 
    rule.style.width?.endsWith('%')
  )
));
```

La larghezza dei tuoi elementi `section` dovrebbe essere impostata usando un valore percentuale.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'section' && 
    rule.style.width?.endsWith('%')
  )
));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

  <header>
    <img src="https://cdn.freecodecamp.org/curriculum/labs/rock.jpg">
    <h1>Pet Rock Festival</h1>
    <p>Join us for the first annual Pet Rock Festival!</p>
    <p><b>When:</b> <time datetime="2025-10-03">October 3rd</time> <b>Where:</b> Boulder Park</p>
  </header>

  <hr>

  <main>
    <section>
      <h2>Festival Highlights</h2>
      <ul>
        <li>Pet Rock Obedience Shows</li>
        <li>Best Dressed Pet Rock Contest</li>
        <li>Pet Rock Uphill Racing</li>
      </ul>
    </section>

    <section>
      <h2>Special Guests</h2>
      <ul>
        <li>John Gravel</li>
        <li>Wendy Cobblestone</li>
      </ul>
    </section>

    <section>
      <h2>Food Trucks</h2>
      <ul>
        <li>Rock 'n' Roll Tacos</li>
        <li>Granite Grill</li>
        <li>Stone Cold Smoothies</li>
      </ul>
    </section>
  </main>

  <hr>

<footer>
  <p>&copy; 2025 Pet Rock Festival</p>
</footer>

</body>

</html>
```

```css
html {
  background-color: gray;
}

body {
  padding: 50px 10px;
  margin: 0 auto;
  width: 90vw;
  min-height: calc(100vh - 100px);
  background-color: white;
  font-family: sans-serif;
  text-align: center;
}

hr {
  margin: 30px auto;
  width: 80%;
  background-color: black;
  height: 2px;
}

main {
  margin: 0 auto;
}

section {
  text-align: center;
  padding: 50px 5px;
  vertical-align: top;
  display: inline-block;
  min-height: 100px;
  width: 30%;
}

ul {
  text-align: left;
  list-style: inside;
  padding-left: 10px;
}
```
