---
id: 67172b43f84bcd2dec238a3d
title: Створіть органайзер для книжок
challengeType: 26
dashedName: build-a-book-organizer
---

# --description--

**Мета:** виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Ви повинні мати масив об’єктів під назвою `books`, де кожен об’єкт масиву має рядки `title` й `authorName`, а також число `releaseYear`.

2. Масив `books` повинен мати принаймні три об’єкти.

3. Ви повинні мати функцію зворотного виклику під назвою `sortByYear`, яка приймає дві книжки як параметри для відсортування масиву.

4. Функція `sortByYear` має повернути `-1`, якщо значення `releaseYear` першої книжки менше за значення другої.

5. Функція `sortByYear` має повернути `1`, якщо значення `releaseYear` першої книжки більше за значення другої.

6. Функція `sortByYear` має повернути `0`, якщо значення `releaseYear` двох книжок однакові.

7. Відфільтруйте книжки, написані після певного року (наприклад, 1950), з масиву `books`, та збережіть відфільтрований масив як новий масив під назвою `filteredBooks`.

8. Відсортуйте книжки в масиві `filteredBooks` за роком видання (`releaseYear`) в порядку зростання. You learned in a prior lecture that the `sort()` method will sort the array in place. Це означає, що масив `filteredBooks` буде змінено.

# --hints--

У коді має бути функція `sortByYear`.

```js
assert.isFunction(sortByYear);
```

Функція `sortByYear` має приймати два параметри.

```js
assert.lengthOf(sortByYear, 2);
```

Функція `sortByYear` має повернути `-1`, якщо значення `releaseYear` об’єкта `book1` менше за значення об’єкта `book2`; `1`, якщо значення `releaseYear` об’єкта `book1` більше за значення об’єкта `book2`; або `0` в інших випадках.

```js
assert.equal(sortByYear({releaseYear: 1913}, {releaseYear: 1925}), -1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1913}), 1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1925}), 0);
```

У коді має бути масив `books`.

```js
assert.isArray(books)
```

Масив `books` має містити принаймні три книжки.

```js
assert.isAtLeast(books.length, 3)
```

Масив `books` має містити об’єкти, кожен з яких має рядки `title` й `authorName`, а також число `releaseYear`.

```js

books.forEach((book) => {
  assert.isObject(book);
  assert.sameMembers(Object.keys(book), ["title", "authorName", "releaseYear"]);
  assert.isString(book.title);
  assert.isNotEmpty(book.title);
  assert.isString(book.authorName);
  assert.isNotEmpty(book.authorName);
  assert.isNumber(book.releaseYear);
})
```

У коді має бути масив `filteredBooks`.

```js
assert.isArray(filteredBooks)
```

Масив `filteredBooks` має містити декілька книжок, відфільтрованих з масиву `books`, і не може бути порожнім.

```js
assert.isBelow(filteredBooks.length, books.length)
assert.isNotEmpty(filteredBooks)
```

Викличте функцію вищого порядку `sort` на масиві `filteredBooks`, передавши функцію зворотного виклику `sortByYear`.

```js
assert.match(__helpers.removeJSComments(code), /\s*filteredBooks\s*\.\s*sort\s*\(\s*sortByYear\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sortByYear(book1, book2) {
  if (book1.releaseYear < book2.releaseYear) {
    return -1;
  }

  if (book1.releaseYear > book2.releaseYear) {
    return 1;
  }

  return 0;
}

const books = [
  {
    title: "The Great Gatsby",
    authorName: "F. Scott Fitzgerald",
    releaseYear: 1925,
  },
  {
    title: "In Search of Lost Time",
    authorName: "Marcel Proust",
    releaseYear: 1913,
  },
  {
    title: "Ulysses",
    authorName: "James Joyce",
    releaseYear: 1922,
  },
  {
    title: "One Hundred Years of Solitude",
    authorName: "Gabriel García Márquez",
    releaseYear: 1967,
  },
  {
    title: "The Catcher in the Rye",
    authorName: "J. D. Salinger",
    releaseYear: 1951,
  },
];

let filteredBooks = books.filter((book) => book.releaseYear < 1950);

filteredBooks.sort(sortByYear);

console.log("Books Written Before 1950 (sorted according to release year)");
filteredBooks.forEach((book) => {
  console.log(`${book.title} by ${book.authorName} (${book.releaseYear})`);
});
```
