---
id: 68c15a2a1f0e70ca7154aba5
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذه الورشة، ستبني نظام تتبع الرواتب للموظفين.

ابدأ بإنشاء صنف باسم `Employee`. بداخله أنشئ طريقة `__init__` بمعلمات `self` و`name` و`level`. داخل طريقة `__init__`، قم بتعيين `name` و`level` إلى الخاصيتين ذات الاسم نفسه في الكائن.

# --hints--

يجب أن يكون لديك صنف باسم `Employee`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Employee")`)) })
```

يجب أن يحتوي صنف `Employee` على طريقة `__init__` بثلاث معلمات هي `self` و`name` و`level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_args("self, name, level")`)) })
```

يجب أن تضبط طريقة `__init__` الخاصية `self.name` لتساوي `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

يجب أن تضبط طريقة `__init__` الخاصية `self.level` لتساوي `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--

```
