---
id: 68c16924498728d259e4781c
title: خطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

أنشئ الآن مثيلًا من `Employee` صنف مع تمرير السلسلتين النصيتين `Charlie Brown` و `trainee`. عيّن المثيل إلى متغير اسمه `charlie_brown`.

# --hints--

يجب أن يكون لديك متغير اسمه `charlie_brown`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("charlie_brown")`)) })
```

يجب أن تعيّن `Employee('Charlie Brown', 'trainee')` إلى `charlie_brown`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("charlie_brown").is_equivalent("charlie_brown = Employee('Charlie Brown', 'trainee')")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level

--fcc-editable-region--

--fcc-editable-region--

```
