---
id: 68c18b1d0fce081f85312d97
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

في درس سابق، تعلمت أن الخاصية التي تبدأ بشرطة سفلية واحدة مخصصة للاستخدام الداخلي حسب الاصطلاح.

غيّر كلًا من الخاصيتين `name` و `level` إلى `_name` و `_level`، إذ لا يُفترض تعديلهما من خارج الصنف.

لاحظ أن هذا لا يمنع الوصول إلى الخاصية أو تعديلها من خارج الصنف. أيضًا، في بايثون توجد دائمًا طريقة للوصول إلى الخصائص الخاصة (التي تبدأ بشرطتين سفليتين) كذلك.

# --hints--

يجب أن تُعيّن الطريقة `__init__` الخاصية `self._name` إلى `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._name = name")`)) })
```

يجب أن تُعيّن الطريقة `__init__` الخاصية `self._level` إلى `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._level = level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
--fcc-editable-region--
        self.name = name
        self.level = level
--fcc-editable-region--

charlie_brown = Employee('Charlie Brown', 'trainee')

```
