---
id: 68c27eb292bfec6b9402d58f
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

أضف طريقة `__str__` إلى صنف `Employee`. اجعلها تُرجع سلسلة نصية بصيغة f-string بالشكل `name: level`، مع استبدال `name` و`level` بالخاصيتين المقابلتين.

بعد ذلك، اطبع `charlie_brown` في وحدة التحكم.

# --hints--

يجب أن يحتوي صنف `Employee` على طريقة `__str__` بمعلمة `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_args("self")`)) })
```

يجب أن تُرجع طريقة `__str__` سلسلة نصية بصيغة f-string بالشكل `name: level`، مع استبدال `name` و`level` بقيم `self._name` و`self._level` على التوالي.

```js
({ test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    empl = Employee("Frank", "dreamer")
    print(empl)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "Frank: dreamer" in output
`) })
```

يجب أن تطبع `charlie_brown` في وحدة التحكم.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown)")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

--fcc-editable-region--
    

charlie_brown = Employee('Charlie Brown', 'trainee')

--fcc-editable-region--

```
