---
id: 68c281cecda94ba2c6838e17
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

الآن بعد أن عرّفت دالة getter لـ `name`، يمكنك الوصول إلى الخاصية `_name` عبر الخاصية `name`. لذا اطبع `charlie_brown.name` في وحدة التحكم.

# --hints--

يجب أن تطبع `charlie_brown.name` في وحدة التحكم.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown.name)")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    @property
    def name(self):
        return self._name

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--

--fcc-editable-region--

```
