---
id: 68c27fa1697ceb7c578b2c9f
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

يُستخدم المزخرف `@property` في Python لتحويل طريقة إلى خاصية. عادةً ما يُستخدم لتعريف طرق getter، وهي طرق تُستخدم لاسترجاع قيمة خاصية:

```py
class Person:
    def __init__(self, name):
        self._name = name

    @property
    def name(self):
        return self._name

p = Person('Alice')
print(p.name)  # Alice
```

أنشئ طريقة باسم `name` تحتوي على معلمة `self` وزخرفها بـ `@property`. داخل الطريقة، أعد `self._name`.
    
# --hints--

يجب أن يحتوي صنف `Employee` على طريقة `name` بمعلمة `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_args("self")`)) })
```

يجب زخرفة طريقة `name` بـ `@property`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_decorators("property")`)) })
```

يجب أن تُعيد طريقة `name` القيمة `self._name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_return("self._name")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

--fcc-editable-region--
    
--fcc-editable-region--

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)

```
