---
id: 68c28d659d345270d4a61084
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

طريقة `__repr__` هي طريقة خاصة من المفترض أن تُرجع تمثيلًا كسلسلة نصية للكائن يمكن استخدامها لإنشائه.

على سبيل المثال، يجب أن تُرجع طريقة `__repr__` الخاصة بـ `Employee('Charlie', 'developer')` السلسلة النصية `Employee('Charlie', 'developer')`، وهي نفس السلسلة المستخدمة لإنشاء الكائن.

أضف إلى صنف `Employee` الخاص بك طريقة `__repr__` بمعلمة `self`، واجعلها تُرجع سلسلة نصية يمكن استخدامها لإنشاء الكائن.

# --hints--

يجب أن يحتوي صنف `Employee` الخاص بك على طريقة `__repr__` بمعلمة `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__repr__").has_args("self")`)) })
```

يجب أن تُرجع طريقة `__repr__` الخاصة بك السلسلة النصية المستخدمة لإنشاء الكائن.

```js
({ test: () => runPython(`
  empl = Employee("Frank", "dreamer")
  assert repr(empl) == 'Employee("Frank", "dreamer")' or repr(empl) == "Employee('Frank', 'dreamer')"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self.name}: {self.level}'
        
--fcc-editable-region--
    
--fcc-editable-region--
    
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)

```
