---
id: 68c2d31a86d5672ca49d4521
title: الخطوة 14
challengeType: 20
dashedName: step-14
---

# --description--

حان الوقت الآن لإضافة بعض التحقق إلى الطريقة `__init__`. في بداية الطريقة، أنشئ جملة `if` تتحقق مما إذا كان `name` أو `level` ليسا من نوع `str`.

داخل جملة `if`، ارفع استثناء `TypeError` مع الرسالة `'name' and 'level' attribute must be of type 'str'.`

# --hints--

يجب أن تحتوي طريقتك `__init__` على جملة `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0]`)) })
```

يجب أن تتحقق جملة `if` مما إذا كان `name` أو `level` ليسا من نوع `str`.

```js
({ test: () => runPython(`
  cond = _Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0].find_conditions()[0]
  conditions = [
    "not (isinstance(name, str) and isinstance(level, str))",
    "not isinstance(name, str) or not isinstance(level, str)",
    "not (isinstance(level, str) and isinstance(name, str))",
    "not isinstance(level, str) or not isinstance(name, str)",
  ]
  assert any(cond.is_equivalent(c) for c in conditions)
`) })
```

يجب أن ترفع جملة `if` استثناء `TypeError` مع الرسالة `'name' and 'level' attribute must be of type 'str'.`

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError("\\'name\\' and \\'level\\' attribute must be of type \\'str\\'.")')`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
--fcc-editable-region--
        
--fcc-editable-region--
        self._name = name
        self._level = level
        
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)

```
