---
id: 68c28c29026ad559f8ab4329
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

اطبع الآن `charlie_brown.level` في وحدة التحكم.

# --hints--

يجب أن تطبع `charlie_brown.level`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown.level)")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level
        
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)
--fcc-editable-region--

--fcc-editable-region--

```
