---
id: 68c28b7cc3ab8f4daf144b5e
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

اتبع ما فعلته في الخطوات السابقة، وأنشئ دالة getter للخاصية `_level`.

# --hints--

يجب أن يحتوي صنف `Employee` على طريقة `level` بمعلمة `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_args("self")`)) })
```

يجب تزيين طريقة `level` بـ `@property`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_decorators("property")`)) })
```

يجب أن تُرجع طريقة `level` القيمة `self._level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_return("self._level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'
    
    @property
    def name(self):
        return self._name

--fcc-editable-region--
    
--fcc-editable-region--

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)

```
