---
id: 68c7f7012a700243eff1cbc0
title: الخطوة 16
challengeType: 20
dashedName: step-16
---

# --description--

يجب اختيار المستوى المستخدم لإنشاء موظف من بين مستويات محددة. ستستخدم `_base_salaries` للتحقق من صحة المستوى وتعيين الراتب المناسب للموظف.

بعد جملة `if` الموجودة لديك، أنشئ جملة `if` أخرى تتحقق مما إذا كان `level` غير موجود في `Employee._base_salaries`.

داخل جملة `if` الجديدة، ارفع استثناء `ValueError` مع الرسالة `Invalid value '{level}' for 'level' attribute.`، حيث يجب استبدال `{level}` بقيمة المعلمة `level`.

# --hints--

عندما يكون `level` غير موجود في `Employee._base_salaries`، يجب أن ترفع استثناء `ValueError` مع الرسالة `Invalid value '{level}' for 'level' attribute.`، حيث يجب استبدال `{level}` بقيمة المعلمة `level`.

```js
({ test: () => runPython(`
  try:
    Employee("Frank", "dreamer")
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
    assert False, "Expected to raise ValueError with invalid level"
`) })
```

لا يجب أن ترفع أي استثناء عندما يكون `level` موجودًا في `Employee._base_salaries`.

```js
({ test: () => runPython(`
  levels = [
    "trainee",
    "junior",
    "mid-level",
    "senior"
  ]
  for level in levels:
    try:
      Employee("Frank", level)
    except Exception:
      assert False, f"Expected not to raise ValueError with level '{level}'"
    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
--fcc-editable-region--
        
--fcc-editable-region--
        self._name = name
        self._level = level
        
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)

```
