---
id: 68c8002beb157fe47d2e7699
title: الخطوة 20
challengeType: 20
dashedName: step-20
---

# --description--

المُعيّن هو طريقة تُستخدم لتعيين قيمة خاصية، مما يسمح بإجراء تحقق وفرض قيود. يمكنك إنشاء مُعيّن باستخدام المزخرف `@propertyName.setter`، حيث يجب أن يتطابق `propertyName` مع اسم الخاصية التي تريد تعيينها:

```py
class Person:
    def __init__(self, name):
        self.name = name  # Calling the setter

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, value):
        self._name = value

p = Person('Alice')
p.name = 'Abigail' # Calls the setter
print(p.name) # Abigail
```

بعد طريقة الجلب الخاصة بك، أنشئ طريقة `name` بمعلمات `self` و `new_name`. زين الطريقة بـ `@name.setter`. داخل الطريقة، عيّن `self._name` إلى `new_name`.

# --hints--

يجب أن يحتوي صنف `Employee` على طريقة `name` بمعلمات `self` و `new_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_args("self, new_name")`)) })
```

يجب أن تُزَيَّن طريقة `name` بـ `@name.setter`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_decorators("name.setter")`)) })
```

يجب أن تُعيّن طريقة `name` القيمة `self._name` إلى `new_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_stmt("self._name = new_name")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

--fcc-editable-region--
    
--fcc-editable-region--

    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')

```
