---
id: 68ca758f8160b11757f877ae
title: الخطوة 32
challengeType: 20
dashedName: step-32
---

# --description--

أخيرًا، أنشئ جملة `if` رابعة ترفع `ValueError` مع الرسالة `Cannot change to lower level.` عندما يكون الراتب الأساسي للمستوى الجديد أقل من الراتب الأساسي للمستوى الحالي.

استخدم `hasattr(self, '_level')` لتجنب `AttributeError` أثناء التهيئة.

# --hints--

يجب أن تحتوي على جملة `if` رابعة داخل محدد الخاصية `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[3]`)) })
```

يجب أن تستخدم جملة `if` الخاصة بك `hasattr(self, '_level')` للتحقق من وجود `_level` قبل المقارنة.

```js
({ test: () => assert(runPython(`"hasattr(self, '_level')" in str(_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[3].find_conditions()[0])`)) })
```

يجب أن تقارن جملة `if` الخاصة بك بين `Employee._base_salaries[new_level]` و `Employee._base_salaries[self.level]` للتحقق مما إذا كان المستوى الجديد لديه راتب أساسي أقل.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[3].find_conditions()[0]
cond.is_equivalent("hasattr(self, '_level') and Employee._base_salaries[new_level] < Employee._base_salaries[self.level]") or cond.is_equivalent("hasattr(self, '_level') and Employee._base_salaries[self.level] > Employee._base_salaries[new_level]")
`)) })
```

عندما يكون `new_level` أقل من `self.level`، يجب أن ترفع `ValueError` مع الرسالة `Cannot change to lower level.`

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'junior')
  try:
    emp.level = "trainee"
  except ValueError as e:
    assert str(e) == "Cannot change to lower level."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

لا يجب أن ترفع أي استثناء عندما يكون `new_level` أعلى من `self.level`.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'junior')
  try:
    emp.level = "senior"
  except Exception:
    assert False, "Expected not to raise exception with valid new_level"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
        if hasattr(self, '_level') and new_level == self.level:
            raise ValueError(f"'{self.level}' is already the selected level.")
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level
    
    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')

```
