---
id: 68e2c945cc1d8e778152be31
title: Побудова інтерфейсу гравця
challengeType: 27
dashedName: build-a-player-interface
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні визначити абстрактний клас з назвою `Player`, який наслідується від класу `abc.ABC`.
1. Клас `Player` має мати метод `__init__`, який встановлює:

   - Атрибут `moves` як порожній список.
   - Атрибут `position` у значення `(0, 0)`.
   - Атрибут `path` як список, що містить початкову позицію.

1. Клас `Player` має мати метод з назвою `make_move`, який:

   - Використовує `random.choice` для отримання випадкового ходу з атрибута `moves` (визначеного в конкретному класі).
   - Додає значення вибраного ходу до поточної позиції та оновлює атрибут `position`.
   - Додає новий кортеж `position` до атрибута `path`.
   - Повертає новий `position`.

1. Клас `Player` має мати абстрактний метод з назвою `level_up`, який потрібно реалізувати в конкретних класах.
1. Ви повинні визначити клас `Pawn`, який наслідується від класу `Player`.
1. Клас `Pawn` має використовувати `super()` для виклику методу `__init__` батьківського елемента, а потім встановити атрибут `moves` як список кортежів, що представляють координати `x, y`.
1. Кожен кортеж координат має представляти рух на `1` одиниць у таких напрямках: вгору, вниз, вліво, вправо.
1. Клас `Pawn` має реалізувати конкретний метод `level_up`, додавши більше ходів до атрибута `moves`. Додані ходи мають представляти чотири діагональні рухи (наприклад, `1` одиниць вниз плюс `1` одиниць вліво).

**Примітка:** Модулі стандартної бібліотеки слід імпортувати без використання псевдонімів. Тести, пов’язані з класом `Player`, не пройдуть, доки клас `Pawn` не стане інстанційованим.

# --hints--

Ви повинні мати клас з назвою `Player`.

```js
({ test: () => runPython(`assert _Node(_code).has_class("Player")`) })
```

Клас `Player` має наслідуватися від класу `ABC` модуля `abc`.

```js
({ test: () => runPython(`
from abc import ABC
import ast
cls = _Node(_code).find_class("Player")
assert cls.inherits_from("ABC") or cls.inherits_from("abc.ABC")
`) })
```

Клас `Player` має мати метод `__init__`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("__init__")`) })
```

Метод `Player` класу `__init__` має мати один параметр `self`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("__init__").has_args("self")`) })
```

Метод `Player` класу `__init__` має встановлювати `self.moves` як порожній список.

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.moves = []") or init.has_stmt("self.moves = list()")
`) })
```

Метод `Player` класу `__init__` має встановлювати `self.position` у значення `(0, 0)`.

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.position = (0, 0)")
`) })
```

Метод `Player` класу `__init__` має встановлювати `self.path` як список, що містить початкову позицію.

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.path = [self.position]")
`) })
```

Клас `Player` має мати метод `make_move`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("make_move")`) })
```

Метод `Player` класу `make_move` має мати один параметр `self`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("make_move").has_args("self")`) })
```

Метод `Player` класу `make_move` має використовувати `random.choice` для отримання випадкового ходу з атрибута `moves`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").block_has_call("choice", "make_move")`) })
```

Метод `Player` класу `make_move` має оновлювати атрибут `position`, додаючи до нього координати випадково вибраного ходу.

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  p.make_move()
  assert p.position == (1, 2)
  p.make_move()
  assert p.position == (2, 4)

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

Метод `Player` класу `make_move` має додавати новий кортеж `position` до атрибута `path`.

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  p.make_move()
  assert p.path == [(0, 0), (1, 2)]
  p.make_move()
  assert p.path == [(0, 0), (1, 2), (2, 4)]

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

Метод `Player` класу `make_move` має повертати оновлений атрибут `position`.

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  assert p.make_move() == (1, 2)

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

Клас `Player` має мати метод `level_up`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("level_up")`) })
```

Метод `Player` класу `level_up` має мати один параметр `self`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("level_up").has_args("self")`) })
```

Метод `Player` класу `level_up` має бути абстрактним.

```js
({ test: () => runPython(`
target = _Node(_code).find_class("Player").find_function("level_up")
assert target.has_decorators("abstractmethod") or target.has_decorators("abc.abstractmethod")
`) })
```

Клас `Player` має бути абстрактним класом.

```js
({ test: () => runPython(`
try:
  Player()
except TypeError as e:
  assert str(e) == "Can't instantiate abstract class Player with abstract method level_up"
else:
  assert False, "Player class should not be instantiable"
`) })
```

Ви повинні мати клас з назвою `Pawn`.

```js
({ test: () => runPython(`assert _Node(_code).has_class("Pawn")`) })
```

Клас `Pawn` має наслідуватися від класу `Player`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").inherits_from("Player")`) })
```

Клас `Pawn` має мати метод `__init__`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").has_function("__init__")`) })
```

Метод `Pawn` класу `__init__` має мати один параметр `self`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").find_function("__init__").has_args("self")`) })
```

Метод `Pawn` класу `__init__` має викликати метод `__init__` батьківського елемента, використовуючи функцію `super`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").find_function("__init__").has_stmt("super().__init__()")`) })
```

Метод `Pawn` класу `__init__` має встановлювати атрибут `moves` як список кортежів, що представляють координати `x, y`, де кожен кортеж координат представляє рух на `1` одиниць у таких напрямках: вгору, вниз, вліво, вправо.

```js
({ test: () => runPython(`
p = Pawn()
moves = [(0, 1), (1, 0), (-1, 0), (0, -1)]
assert isinstance(p.moves, list)
assert len(p.moves) == 4
assert all(move in p.moves for move in moves)
`) })
```

Клас `Pawn` має мати метод `level_up`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").has_function("level_up")`) })
```

Метод `Pawn` класу `level_up` має додати чотири діагональні рухи на `1` одиниць до атрибута `moves`.

```js
({ test: () => runPython(`
p = Pawn()
moves = [(0, 1), (1, 0), (-1, 0), (0, -1), (1, 1), (1, -1), (-1, 1), (-1, -1)]
p.level_up()
assert isinstance(p.moves, list)
assert len(p.moves) == 8
assert all(move in p.moves for move in moves)
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
from abc import ABC, abstractmethod
from random import choice


class Player(ABC):
    def __init__(self):
        self.moves = []
        self.position = (0, 0)
        self.path = [self.position]

    def make_move(self):
        x, y = self.position
        move = choice(self.moves)
        new_x = move[0] + x
        new_y = move[1] + y
        self.position = (new_x, new_y)
        self.path.append(self.position)
        return self.position

    @abstractmethod
    def level_up(self):
        pass


class Pawn(Player):
    def __init__(self):
        super().__init__()
        self.moves = [(0, 1), (1, 0), (-1, 0), (0, -1)]

    def level_up(self):
        self.moves.extend([(1, 1), (1, -1), (-1, 1), (-1, -1)])
```
