---
id: 587d78af367417b2b2512b03
title: بناء نموذج استبيان
challengeType: 25
dashedName: build-a-survey-form
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك عنوان صفحة في عنصر `h1` مع `id` يساوي `title`.
1. يجب أن يكون لديك شرح قصير في عنصر `p` مع `id` يساوي `description`.
1. يجب أن يكون لديك عنصر `form` مع `id` يساوي `survey-form`.
1. داخل عنصر النموذج يجب أن يكون لديك حقل `input` مطلوب لإدخال اسمك مع `id` يساوي `name` و`type` يساوي `text`.
1. داخل عنصر النموذج يجب أن يكون لديك حقل `input` مطلوب لإدخال بريدك الإلكتروني مع `id` يساوي `email`.
1. إذا أدخلت بريدًا إلكترونيًا غير منسق بشكل صحيح، يجب أن ترى خطأ تحقق HTML5.
1. داخل عنصر النموذج يجب أن يكون لديك حقل `input` لإدخال رقم مع `id` يساوي `number`.
1. يجب ألا يقبل حقل الرقم إدخال غير الأرقام، إما بمنعك من كتابتها أو بإظهار خطأ تحقق HTML5 (حسب المتصفح).
1. إذا أدخلت أرقامًا خارج نطاق حقل الرقم، المحدد بواسطة خاصيتي `min` و`max`، يجب أن ترى خطأ تحقق HTML5.
1. يجب أن يكون لديك عناصر `label` مقابلة لحقول الإدخال الخاصة بالاسم والبريد الإلكتروني والرقم في النموذج، تصف غرض كل حقل مع المعرفات التالية: `name-label`، `email-label`، و`number-label`.
1. يجب أن يكون لكل من حقول الاسم والبريد الإلكتروني والرقم نص عنصر نائب يصف أو يوجه كل حقل.
1. داخل عنصر النموذج يجب أن يكون لديك عنصر قائمة منسدلة `select` مع `id` يساوي `dropdown` وخياران على الأقل للاختيار.
1. داخل عنصر النموذج يمكنك اختيار خيار من مجموعة تحتوي على زري اختيار على الأقل، مجمعة باستخدام خاصية `name`.
1. داخل عنصر النموذج يمكنك اختيار عدة حقول من سلسلة مربعات اختيار، يجب أن يحتوي كل منها على خاصية `value`.
1. داخل عنصر النموذج يجب أن يكون لديك عنصر `textarea` لتعليقات إضافية.
1. داخل عنصر النموذج يجب أن يكون لديك زر مع `id` يساوي `submit` لإرسال جميع المدخلات.

# --hints--

يجب أن يكون لديك عنصر `h1` مع `id` يساوي `title`.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

يجب ألا يكون `#title` فارغًا.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

يجب أن يكون لديك عنصر `p` مع `id` يساوي `description`.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

يجب ألا يكون `#description` فارغًا.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

يجب أن يكون لديك عنصر `form` مع `id` يساوي `survey-form`.

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

يجب أن يكون لديك عنصر `input` مع `id` يساوي `name`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

يجب أن يكون لـ `#name` خاصية `type` بقيمة `text`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

يجب أن يكون حقل `#name` مطلوب الإدخال.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

يجب أن يكون `#name` عنصرًا تابعًا لـ `#survey-form`.

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

يجب أن يكون لديك عنصر `input` مع `id` يساوي `email`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

يجب أن يكون لـ `#email` خاصية `type` بقيمة `email`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

يجب أن يكون حقل `#email` مطلوب الإدخال.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

يجب أن يكون `#email` عنصرًا تابعًا لـ `#survey-form`.

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

يجب أن يكون لديك عنصر `input` مع `id` يساوي `number`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

يجب أن يكون `#number` عنصرًا تابعًا لـ `#survey-form`.

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

يجب أن يكون لـ `#number` خاصية `type` بقيمة `number`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

يجب أن يحتوي `#number` على خاصية `min` بقيمة رقمية.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

يجب أن يحتوي `#number` على خاصية `max` بقيمة رقمية.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

يجب أن يكون لديك عنصر `label` مع `id` يساوي `name-label`.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

يجب أن يكون لديك عنصر `label` مع `id` يساوي `email-label`.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

يجب أن يكون لديك عنصر `label` مع `id` يساوي `number-label`.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

يجب أن يحتوي `#name-label` على نص يصف حقل الإدخال.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

يجب أن يحتوي `#email-label` على نص يصف حقل الإدخال.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

يجب أن يحتوي `#number-label` على نص يصف حقل الإدخال.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

يجب أن يكون `#name-label` عنصرًا تابعًا لـ `#survey-form`.

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

يجب أن يكون `#email-label` عنصرًا تابعًا لـ `#survey-form`.

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

يجب أن يكون `#number-label` عنصرًا تابعًا لـ `#survey-form`.

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

يجب أن يحتوي `#name` على خاصية `placeholder` وقيمة.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

يجب أن يحتوي `#email` على خاصية `placeholder` وقيمة.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

يجب أن يحتوي `#number` على خاصية `placeholder` وقيمة.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

يجب أن يكون لديك حقل `select` مع `id` يساوي `dropdown`.

```js
const el = document.getElementById('dropdown');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'SELECT');
```

يجب أن يحتوي `#dropdown` على عنصرين `option` قابلين للاختيار على الأقل (غير معطّلين).

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert.isAtLeast(els.length, 2);
```

يجب أن يكون `#dropdown` عنصرًا تابعًا لـ `#survey-form`.

```js
const el = document.querySelector('#survey-form #dropdown');
assert.isNotNull(el);
```

يجب أن يكون لديك على الأقل عنصران من نوع `input` مع `type` يساوي `radio` (أزرار اختيار).

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

يجب أن يكون لديك على الأقل زران اختيار عنصران تابعان لـ `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

يجب أن تحتوي جميع أزرار الاختيار على خاصية `value` وقيمة.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][value=""], input[type="radio"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

يجب أن تحتوي جميع أزرار الاختيار على خاصية `name` وقيمة.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][name=""], input[type="radio"]:not([name])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

يجب أن تحتوي كل مجموعة أزرار اختيار على زرين اختيار على الأقل.

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  });
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert.fail();
});

assert.isAbove(groupKeys.length, 0);
```

يجب أن يكون لديك على الأقل عنصران من نوع `input` مع `type` يساوي `checkbox` (مربعات اختيار) كعناصر تابعة لـ `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert.isAtLeast(els.length, 2);
```

يجب أن تحتوي جميع مربعات الاختيار داخل `#survey-form` على خاصية `value` وقيمة.

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll(
  '#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

يجب أن يكون لديك عنصر `textarea` واحد على الأقل تابع لـ `#survey-form`.

```js
const el = document.querySelector('#survey-form textarea');
assert.isNotNull(el);
```

يجب أن يكون لديك عنصر `input` أو `button` مع `id` يساوي `submit`.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.isTrue(el.tagName === 'INPUT' || el.tagName === 'BUTTON');
```

يجب أن يكون لـ `#submit` خاصية `type` بقيمة `submit`.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.strictEqual(el.type, 'submit');
```

يجب أن يكون `#submit` عنصرًا تابعًا لـ `#survey-form`.

```js
const el = document.querySelector('#survey-form #submit');
assert.isNotNull(el);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Survey Form</title>
  </head>

  <body></body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="container">
      <header class="header">
        <h1 id="title" class="text-center">freeCodeCamp Survey Form</h1>
        <p id="description" class="description text-center">
          Thank you for taking the time to help us improve the platform
        </p>
      </header>
      <form id="survey-form">
        <div class="form-group">
          <label id="name-label" for="name">Name</label>
          <input
            type="text"
            name="name"
            id="name"
            class="form-control"
            placeholder="Enter your name"
            required
          />
        </div>
        <div class="form-group">
          <label id="email-label" for="email">Email</label>
          <input
            type="email"
            name="email"
            id="email"
            class="form-control"
            placeholder="Enter your Email"
            required
          />
        </div>
        <div class="form-group">
          <label id="number-label" for="number"
            >Age<span class="clue">(optional)</span></label
          >
          <input
            type="number"
            name="age"
            id="number"
            min="10"
            max="99"
            class="form-control"
            placeholder="Age"
          />
        </div>
        <div class="form-group">
          <p>Which option best describes your current role?</p>
          <select id="dropdown" name="role" class="form-control" required>
            <option disabled selected value>Select current role</option>
            <option value="student">Student</option>
            <option value="job">Full Time Job</option>
            <option value="learner">Full Time Learner</option>
            <option value="preferNo">Prefer not to say</option>
            <option value="other">Other</option>
          </select>
        </div>

        <div class="form-group">
          <p>Would you recommend freeCodeCamp to a friend?</p>
          <label>
            <input
              name="user-recommend"
              value="definitely"
              type="radio"
              class="input-radio"
              checked
            />Definitely</label
          >
          <label>
            <input
              name="user-recommend"
              value="maybe"
              type="radio"
              class="input-radio"
            />Maybe</label
          >

          <label
            ><input
              name="user-recommend"
              value="not-sure"
              type="radio"
              class="input-radio"
            />Not sure</label
          >
        </div>

        <div class="form-group">
          <p>What is your favorite feature of freeCodeCamp?</p>
          <select id="most-like" name="mostLike" class="form-control" required>
            <option disabled selected value>Select an option</option>
            <option value="challenges">Challenges</option>
            <option value="projects">Projects</option>
            <option value="community">Community</option>
            <option value="openSource">Open Source</option>
          </select>
        </div>

        <div class="form-group">
          <p>
            What would you like to see improved?
            <span class="clue">(Check all that apply)</span>
          </p>

          <label
            ><input
              name="prefer"
              value="front-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Front-end Projects</label
          >
          <label>
            <input
              name="prefer"
              value="back-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Back-end Projects</label
          >
          <label
            ><input
              name="prefer"
              value="data-visualization"
              type="checkbox"
              class="input-checkbox"
            />Data Visualization</label
          >
          <label
            ><input
              name="prefer"
              value="challenges"
              type="checkbox"
              class="input-checkbox"
            />Challenges</label
          >
          <label
            ><input
              name="prefer"
              value="open-source-community"
              type="checkbox"
              class="input-checkbox"
            />Open Source Community</label
          >
          <label
            ><input
              name="prefer"
              value="gitter-help-rooms"
              type="checkbox"
              class="input-checkbox"
            />Gitter help rooms</label
          >
          <label
            ><input
              name="prefer"
              value="videos"
              type="checkbox"
              class="input-checkbox"
            />Videos</label
          >
          <label
            ><input
              name="prefer"
              value="city-meetups"
              type="checkbox"
              class="input-checkbox"
            />City Meetups</label
          >
          <label
            ><input
              name="prefer"
              value="wiki"
              type="checkbox"
              class="input-checkbox"
            />Wiki</label
          >
          <label
            ><input
              name="prefer"
              value="forum"
              type="checkbox"
              class="input-checkbox"
            />Forum</label
          >
          <label
            ><input
              name="prefer"
              value="additional-courses"
              type="checkbox"
              class="input-checkbox"
            />Additional Courses</label
          >
        </div>

        <div class="form-group">
          <p>Any comments or suggestions?</p>
          <textarea
            id="comments"
            class="input-textarea"
            name="comment"
            placeholder="Enter your comment here..."
          ></textarea>
        </div>

        <div class="form-group">
          <button type="submit" id="submit" class="submit-button">
            Submit
          </button>
        </div>
      </form>
    </div>
  </body>
</html>
```
