---
id: 698b42492a7f5f737895892e
title: Erstellen Sie einen Konverter von Adjazenzliste zu Matrix
challengeType: 26
dashedName: build-an-adjacency-list-to-matrix-converter
---

# --description--

In diesem Labor erstellen Sie eine Funktion, die eine Adjazenzlisten-Darstellung eines Graphen in eine Adjazenzmatrix umwandelt. Eine Adjazenzliste ist ein Objekt, bei dem jeder Schlüssel einen Knoten darstellt und der zugehörige Wert ein Array von Knoten ist, mit denen der Schlüssel-Knoten verbunden ist. Eine Adjazenzmatrix ist ein 2D-Array, bei dem der Eintrag an der Position `[i][j]` `1` ist, wenn eine Kante von Knoten `i` zu Knoten `j` existiert, und `0` sonst.

Zum Beispiel, gegeben die Adjazenzliste:

```js
{
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
}
```

Die entsprechende Adjazenzmatrix wäre:

```js
[
  [0, 1, 1, 0],
  [0, 0, 1, 0],
  [1, 0, 0, 1],
  [0, 0, 1, 0]
]
```

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `adjacencyListToMatrix` definieren, um eine Adjazenzliste in eine Adjazenzmatrix umzuwandeln.
2. Die Funktion sollte ein Objekt als Argument nehmen, das die Adjazenzliste eines ungewichteten (entweder ungerichteten oder gerichteten) Graphen darstellt.
3. Die Funktion sollte:
   - Die Adjazenzliste in eine Adjazenzmatrix umwandeln.
   - Jede Zeile der Adjazenzmatrix ausgeben.
   - Die Adjazenzmatrix zurückgeben.

Zum Beispiel sollte `adjacencyListToMatrix({0: [2], 1: [2, 3], 2: [0, 1, 3], 3: [1, 2]})` ausgeben:

```md
[0, 0, 1, 0]
[0, 0, 1, 1]
[1, 1, 0, 1]
[0, 1, 1, 0]
```

und `[[0, 0, 1, 0], [0, 0, 1, 1], [1, 1, 0, 1], [0, 1, 1, 0]]` zurückgeben.

# --hints--

Sie sollten eine Funktion mit dem Namen `adjacencyListToMatrix` definieren.

```js
assert.isFunction(adjacencyListToMatrix);
```

Die `adjacencyListToMatrix`-Funktion sollte einen Parameter haben.

```js
assert.lengthOf(adjacencyListToMatrix, 1);
```

Die Funktion sollte die Anzahl der Knoten aus der Adjazenzliste korrekt bestimmen.

```js
const adjList = {0: [1], 1: [0], 2: []};
const result = adjacencyListToMatrix(adjList);
assert.lengthOf(result, 3);
result.forEach(arr => assert.lengthOf(arr, 3));
```

Die Funktion sollte die Matrixwerte für vorhandene Kanten korrekt auf `1` setzen.

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

Die Funktion sollte jede Zeile der Matrix ausgeben.

```js
const spy = __helpers.spyOn(console, "log");
try {
  const adjList = {0: [1], 1: []}
  adjacencyListToMatrix(adjList);
  const args = spy.calls;
  assert.sameDeepOrderedMembers(args[0], [ [0, 1] ]);
  assert.sameDeepOrderedMembers(args[1], [ [0, 0] ]);
} catch (e) {
  assert.fail(e);
} finally {
  spy.restore();
}
```

Die Funktion sollte die Adjazenzmatrix zurückgeben.

```js
const adjList = {0: [1], 1: []};
const result = adjacencyListToMatrix(adjList);
assert.sameDeepOrderedMembers(result, [[0, 1], [0, 0]]);
```

Wenn die Funktion die Adjazenzliste `{0: [1, 2], 1: [2], 2: [0, 3], 3: [2]}` erhält, sollte sie `[[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]]` zurückgeben.

```js
const adjList = {0: [1, 2], 1: [2], 2: [0, 3], 3: [2]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

Wenn die Funktion die Adjazenzliste `{0: [1], 1: [0]}` erhält, sollte sie `[[0, 1], [1, 0]]` zurückgeben.

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

Wenn die Funktion die Adjazenzliste `{0: [], 1: [], 2: []}` erhält, sollte sie `[[0, 0, 0], [0, 0, 0], [0, 0, 0]]` zurückgeben.

```js
const adjList = {0: [], 1: [], 2: []};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 0, 0], [0, 0, 0], [0, 0, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function adjacencyListToMatrix(adjList) {
  const n = Object.keys(adjList).length;

  const adjMatrix = Array.from({ length: n }, () => Array(n).fill(0));

  for (const [srcStr, neighbors] of Object.entries(adjList)) {
    const srcNode = Number(srcStr);
    for (const destNode of neighbors) {
      adjMatrix[srcNode][destNode] = 1;
    }
  }

  for (const row of adjMatrix) {
    console.log(row);
  }

  return adjMatrix;
}

const adjList = {
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
};

const matrix = adjacencyListToMatrix(adjList);
```
