---
id: 669fdfc9a5e3da42d2376609
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

Ora che hai aggiunto la sezione about, prova a cliccare sul collegamento `About` per vedere la pagina saltare a quella sezione.

La sezione successiva nella pagina del blog sarà un elenco di post che parlano di Mr. Whiskers.

Aggiungi un altro elemento `section` con un insieme `id` impostato su `"posts"`.

All'interno dell'elemento `section`, aggiungi un elemento `h2` con il testo `Posts`.

# --hints--

Il tuo codice dovrebbe avere un tag di apertura `section`.

```js
assert.match(code, /<section.*>/i);
```

Il tuo codice dovrebbe avere un tag di chiusura `section`.

```js
assert.match(code, /<\/section>/i);
```

Il tuo elemento `section` dovrebbe avere un insieme `id` impostato su `"posts"`.

```js
assert.match(code, /<section.*id="posts".*>/i);
```

Il tuo elemento `section` dovrebbe essere all'interno del tuo elemento `main`.

```js
const sectionElement = document.getElementById('posts');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

Il tuo elemento `h2` dovrebbe essere dentro l'elemento `section`.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

Il tuo elemento `h2` dovrebbe avere il testo `Posts`.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.innerText.trim(), 'Posts');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </main>
  </body>
</html>
```
