---
id: 669fde2081f65141ad703fe4
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

La prima sezione della pagina sarà la sezione about. La sezione presenterà Mr. Whiskers e darà agli utenti un’idea di cosa tratta questo blog.

All’interno del tuo elemento `main`, aggiungi un elemento `section` con l’attributo `id` impostato su `"about"`.

All’interno dell’elemento `section`, aggiungi un `h2` con il testo `About`.

# --hints--

Il tuo elemento `section` dovrebbe avere un tag di apertura.

```js
assert.match(code, /<section.*>/);
```

Il tuo elemento `section` dovrebbe avere un tag di chiusura.

```js
assert.match(code, /<\/section\>/);
```

Il tuo elemento `section` dovrebbe avere un `id` impostato su `"about"`.

```js
const sectionElement = document.querySelector('section');
assert.strictEqual(sectionElement?.getAttribute("id"), "about");
```

Il tuo elemento `section` dovrebbe essere all’interno del tuo elemento `main`.

```js
const sectionElement = document.getElementById('about');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

Il tuo elemento `h2` dovrebbe avere un tag di apertura.

```js
assert.match(code, /<h2\>/);
```

Il tuo elemento `h2` dovrebbe avere un tag di chiusura.

```js
assert.match(code, /<\/h2\>/);
```

Il tuo elemento `h2` dovrebbe avere il testo `About`. Controlla bene l’ortografia.

```js
const h2Element = document.querySelector('body h2');
assert.strictEqual(h2Element?.innerText.trim(), "About");
```

Il tuo elemento `h2` dovrebbe essere all’interno del tuo elemento `section`.

```js
const h2Element = document.querySelector('#about h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
