---
id: 669fd6fd12918e3de87854d4
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

Per il tuo blog, dovrebbe esserci un modo per gli utenti di navigare tra le diverse sezioni della pagina.

L'elemento `nav` viene usato per fornire collegamenti di navigazione ad altre sezioni nel documento o ad altre sezioni del sito web. Spesso vedrai l'elemento `nav` usato per menu o indici.

Ecco un esempio di utilizzo dell'elemento `nav`:

```html
<nav>
  <ul>
    <li><a href="#">Home</a></li>
    <li><a href="#about">About</a></li>
    <li><a href="#contact">Contact</a></li>
  </ul>
</nav>
```

Sotto il tuo elemento `figure`, aggiungi un elemento `nav` con un elemento `ul` annidato all'interno.

All'interno dell'elemento `ul`, aggiungi tre elementi `li`.

# --hints--

Il tuo elemento `nav` dovrebbe avere un tag di apertura.

```js
assert.match(code, /<nav\>/);
```

Il tuo elemento `nav` dovrebbe avere un tag di chiusura.

```js
assert.match(code, /<\/nav\>/);
```

Il tuo elemento `nav` dovrebbe essere all'interno del tuo elemento `header`.

```js
const navElement = document.querySelector('nav');
assert.strictEqual(navElement?.parentElement.tagName, 'HEADER');
```

Il tuo elemento `ul` dovrebbe avere un tag di apertura.

```js
assert.match(code, /<ul\>/);
```

Il tuo elemento `ul` dovrebbe avere un tag di chiusura.

```js
assert.match(code, /<\/ul\>/);
```

Il tuo elemento `ul` dovrebbe essere all'interno del tuo elemento `nav`.

```js
const ulElement = document.querySelector('ul');
assert.strictEqual(ulElement?.parentElement.tagName, 'NAV');
```

Il tuo elemento `ul` dovrebbe contenere tre elementi `li`.

```js
const ulElement = document.querySelector("ul");

const liElements = ulElement?.querySelectorAll("li");
assert.strictEqual(liElements?.length, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
