---
id: 669fcb5106eeb13ab7afcd09
title: Passaggio 1
challengeType: 0
dashedName: step-1
---

# --description--

In questo workshop, praticherai l’uso dell’HTML semantico costruendo una pagina blog dedicata a Mr. Whiskers, il gatto.

La prima sezione che costruirai è l’<dfn>intestazione</dfn> della pagina.

L’elemento `header` viene usato per rappresentare contenuti introduttivi come la navigazione della pagina e altre informazioni iniziali.

Ecco un esempio che usa l’elemento `header`:

```html
<header>
  <h1>Main Page Title Goes Here</h1>
  <img src="example-logo.png" alt="Example logo" />
</header>
```

All’interno dell’elemento `body`, aggiungi un elemento `header`.

# --hints--

Dovresti avere un tag di apertura `header`.

```js
assert.match(code, /<header>/i);
```

Dovresti avere un tag di chiusura `header`.

```js
assert.match(code, /<\/header>/i);
```

Il tuo elemento `header` dovrebbe essere all’interno dell’elemento `body`.

```js
assert.exists(document.querySelector('body > header'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </body>
</html>
```
