---
id: 669fd003cb89ee3c2402e041
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

In questo contenuto introduttivo, dovrai mostrare un'immagine di Mr. Whiskers con un sottotitolo.

Sotto l'elemento `h1`, inizia aggiungendo un elemento `figure`.

All'interno dell'elemento `figure`, aggiungi un elemento `img`.

L'attributo `src` dell'elemento `img` dovrebbe avere un valore `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"` e il testo `alt` dovrebbe avere un valore `"a cat in the garden"`.

Sotto il tuo elemento `img`, aggiungi un `figcaption` con il testo `Mr. Whiskers in the Garden`.

# --hints--

Il tuo elemento `figure` dovrebbe avere un tag di apertura. 

```js
assert.match(code, /<figure\>/);
```

Il tuo elemento `figure` dovrebbe avere un tag di chiusura. 

```js
assert.match(code, /<\/figure\>/);
```

Il tuo elemento `figure` dovrebbe essere all'interno del tuo elemento `header`.

```js
const figureElement = document.querySelector('figure');
assert.equal(figureElement?.parentElement.tagName, 'HEADER');
```

Dovresti avere un elemento `img`.

```js
assert.exists(document.querySelector("img"));
```

Il tuo elemento `img` dovrebbe essere all'interno del tuo elemento `figure`.

```js
const imgElement = document.querySelector('img');
assert.strictEqual(imgElement?.parentElement.tagName, 'FIGURE');
```

Il tuo elemento `img` dovrebbe avere un attributo `src` impostato su `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`.

```js
const imgElement = document.querySelector('img');
const src = imgElement?.getAttribute('src');

assert.strictEqual(src, 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Il tuo elemento `img` dovrebbe avere un attributo `alt` impostato su `"a cat in the garden"`.

```js
const imgElement = document.querySelector('img');
const alt = imgElement?.getAttribute('alt');

assert.strictEqual(alt, 'a cat in the garden');
```

Il tuo elemento `figcaption` dovrebbe avere un tag di apertura. 

```js
assert.match(code, /<figcaption\>/);
```

Il tuo elemento `figcaption` dovrebbe avere un tag di chiusura. 

```js
assert.match(code, /<\/figcaption\>/);
```

Il tuo elemento `figcaption` dovrebbe essere all'interno del tuo elemento `figure`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.parentElement.tagName, 'FIGURE');
```

Il tuo elemento `figcaption` dovrebbe avere il testo `Mr. Whiskers in the Garden`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.innerText.trim(), "Mr. Whiskers in the Garden")
```

Il tuo `figcaption` dovrebbe essere sotto l'elemento `img`. Li hai nell'ordine sbagliato.

```js
const figureElement = document.querySelector('figure');
assert.strictEqual(figureElement?.lastElementChild?.tagName, "FIGCAPTION")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
