---
id: 669fdd3965f36f40af9615e5
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

Ora che hai finito di costruire l'intestazione della pagina, devi iniziare ad aggiungere il contenuto principale.

Sotto il tuo `header`, aggiungi un elemento `main`.

# --hints--

Il tuo elemento `main` dovrebbe avere un tag di apertura.

```js
assert.match(code, /<main\>/);
```

Il tuo elemento `main` dovrebbe avere un tag di chiusura.

```js
assert.match(code, /<\/main\>/);
```

Il tuo elemento `main` dovrebbe essere sotto il tuo elemento `header`.

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
