---
id: 669fdc11c9b0263fe0814a7a
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

All'interno di ciascuno degli elementi `li` dovrai inserire un elemento àncora.

Per il primo elemento àncora, il testo dovrebbe essere `About` e il valore dell'attributo `href` dovrebbe essere `"#about"`. Il simbolo cancelletto davanti a `about` rappresenta un nome `id`, che verrà aggiunto più avanti nel progetto.

Per il secondo elemento àncora, il testo dovrebbe essere `Posts` e il valore dell'attributo `href` dovrebbe essere `"#posts"`.

Per il terzo elemento àncora, il testo dovrebbe essere `Contact` e il valore dell'attributo `href` dovrebbe essere `"#contact"`.

# --hints--

Il tuo primo elemento àncora dovrebbe avere il testo `About` all'interno del tuo primo elemento `li`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "About");
```

Il tuo primo elemento àncora dovrebbe avere un attributo `href` impostato su `"#about"`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#about");
```

Il tuo secondo elemento àncora dovrebbe avere il testo `Posts` all'interno del tuo secondo elemento `li`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
assert.strictEqual(anchorElement?.innerText.trim(), "Posts");
```

Il tuo secondo elemento àncora dovrebbe avere un attributo `href` impostato su `"#posts"`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#posts");
```

Il tuo terzo elemento àncora dovrebbe avere il testo `Contact` all'interno del tuo terzo elemento `li`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "Contact");
```

Il tuo terzo elemento àncora dovrebbe avere un attributo `href` impostato su `"#contact"`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#contact");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
--fcc-editable-region--
          <li></li>
          <li></li>
          <li></li>
--fcc-editable-region--
        </ul>
      </nav>
    </header>
  </body>
</html>
```
