---
id: 66a334a1a7cca6354999f9bf
title: Passaggio 10
challengeType: 0
dashedName: step-10
---

# --description--

Per il primo post del blog, userai un elemento <dfn>article</dfn>.

L'elemento `article` rappresenta contenuto autonomo in una pagina web.

```html
<article>
  <h1>Example heading</h1>
  <p>Example article text</p>
</article>
```

Sotto l'elemento `h2`, aggiungi un elemento `article`.

All'interno dell'elemento `article`, aggiungi un elemento `h3` con il testo `Mr. Whiskers' First Day Home`.

Il motivo per cui qui si usa un `h3` è che mantenere una gerarchia strutturale corretta per gli elementi di intestazione è importante. Poiché il sottotitolo del post è un elemento `h2`, il livello successivo nella gerarchia sarebbe un `h3`.

# --hints--

Il tuo codice dovrebbe avere un tag di apertura `article`.

```js
assert.match(code, /<article>/i);
```

Il tuo codice dovrebbe avere un tag di chiusura `article`.

```js
assert.match(code, /<\/article>/i);
```

Il tuo elemento `article` dovrebbe essere all'interno del tuo elemento `section` con la prop `id` impostata su `posts`.

```js
const articleElement = document.querySelector('#posts article');
assert.strictEqual(articleElement?.parentElement.tagName, 'SECTION');
```

Il tuo codice dovrebbe avere un tag di apertura `h3`.

```js
assert.match(code, /<h3>/i);
```

Il tuo codice dovrebbe avere un tag di chiusura `h3`.

```js
assert.match(code, /<\/h3>/i);
```

Il tuo elemento `h3` dovrebbe avere il testo `Mr. Whiskers' First Day Home`.

```js
const h3Element = document.querySelector('h3');
assert.strictEqual(h3Element?.innerText.trim(), "Mr. Whiskers' First Day Home");
```

Il tuo elemento `h3` dovrebbe essere all'interno del tuo elemento `article`.

```js
const h3Element = document.querySelector('#posts article h3');
assert.strictEqual(h3Element?.parentElement.tagName, 'ARTICLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
