---
id: 66a49f685961e997e337cab1
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Il `header` sarà responsabile di mostrare il titolo principale, l'immagine e la navigazione della pagina per il blog.

All'interno dell'elemento `header`, aggiungi un `h1` con il testo `Welcome to Mr. Whiskers' Blog Page!`.

# --hints--

Dovresti avere un tag di apertura `h1`.

```js
assert.match(code, /<h1>/i);
```

Dovresti avere un tag di chiusura `h1`.

```js
assert.match(code, /<\/h1>/i);
```

Il tuo elemento `h1` dovrebbe contenere il testo `Welcome to Mr. Whiskers' Blog Page!`. Controlla bene che non ci siano errori di ortografia.

```js
const h1Text = document.querySelector("h1")?.innerText.trim();
assert.strictEqual(h1Text, "Welcome to Mr. Whiskers' Blog Page!");
```

Il tuo elemento `h1` dovrebbe essere annidato all'interno del tuo elemento `header`.

```js
const h1Element = document.querySelector("h1");
assert.strictEqual(h1Element?.parentElement.tagName, 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
