---
id: 68c2fb153760d075321c5da3
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Camperbot이 코딩 여정 블로그 페이지를 만들었지만, 이 페이지에는 접근성 문제와 좋지 않은 모범 사례가 있습니다. 이 워크숍에서 해야 할 일은 Camperbot의 문제를 고치는 것입니다.

첫 번째 오류 집합은 제목 사용과 관련이 있습니다. 이전 수업에서 한 페이지에 `h1`는 하나만 사용하는 것이 모범 사례임을 배웠습니다. 그런데 Camperbot은 `h1`를 여러 개 사용하고 있습니다.

`<h1>Welcome to Camperbot's Blog</h1>`는 그대로 두고, 다른 `h1` 요소들을 `h2` 요소로 바꾸세요.

# --hints--

`<h1>Welcome to Camperbot's Blog</h1>`는 변경하지 마세요.

```js
const heading = document.querySelectorAll("h1")[0];
assert.equal(heading?.textContent, "Welcome to Camperbot's Blog");
```

텍스트가 `h1`인 `My Journey Learning to Code` 요소를 `h2` 요소로 바꾸세요.

```js
const firstPost = document.getElementById("post1");
assert.equal(firstPost?.tagName, "H2");
assert.equal(firstPost?.textContent, "My Journey Learning to Code");
```

텍스트가 `h1`인 `Accessibility Matters` 요소를 `h2` 요소로 바꾸세요.

```js
const secondPost = document.getElementById("post2");
assert.equal(secondPost?.tagName, "H2");
assert.equal(secondPost?.textContent, "Accessibility Matters");
```

텍스트가 `h1`인 `What's Next?` 요소를 `h2` 요소로 바꾸세요.

```js
const thirdPost = document.getElementById("post3");
assert.equal(thirdPost?.tagName, "H2");
assert.equal(thirdPost?.textContent, "What's Next?");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Camperbot's Blog</title>
</head>
<body>
  <h1>Welcome to Camperbot's Blog</h1>

  <h2>Navigation</h2>
  <ul>
    <li><a href="#post1">My Journey</a></li>
    <li><a href="#post2">Accessibility</a></li>
    <li><a href="#post3">Next Steps</a></li>
  </ul>
--fcc-editable-region--
  <div>
    <h1 id="post1">My Journey Learning to Code</h1>
    <p>I started learning to code a few months ago and it's been a wild ride!</p>
    
    <h4>Early Challenges</h4>
    <p>At first, syntax was really confusing.</p>
    
    <h4>Breakthroughs</h4>
    <p>Eventually things started to click.</p>
  </div>

  <div>
    <h1 id="post2">Accessibility Matters</h1>
    <p>Today I learned that not everyone uses the web the same way I do.</p>

    <h5>Screen Readers</h5>
    <p>These tools help visually impaired users browse websites.</p>
  </div>

  <div>
    <h1 id="post3">What's Next?</h1>
    <p>I'm excited to dive into JavaScript and build interactive features!</p>

    <h3>Coming soon: My first JavaScript project!</h3>
    <p>Stay tuned for some exciting interactive blog features.</p>
  </div>
--fcc-editable-region--
  <h3>Contact Me</h3>
  <p>Email me at camperbot@blog.io</p>
</body>
</html>
```
