---
id: 68c31dd917c3ae855e8728c7
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

워크숍의 마지막 부분에서는 하단의 연락처 섹션에 몇 가지 변경이 필요합니다.

첫 번째 변경 사항은 연락처 섹션을 `footer` 요소 안에 감싸는 것입니다. 이렇게 하면 화면 읽기 프로그램이 이 섹션을 페이지의 푸터로 인식하는 데 도움이 됩니다.

# --hints--

페이지에 `footer` 요소가 있어야 합니다.

```js
const footerElement = document.querySelector("footer");
assert.isNotNull(footerElement);
```

`footer` 요소 안에는 텍스트 `h3`가 포함된 `Contact Me` 요소가 있어야 합니다.

```js
const h3Element = document.querySelector("footer h3");
assert.isNotNull(h3Element);
assert.strictEqual(h3Element?.textContent, "Contact Me");
```

`footer` 요소 안에는 텍스트 `p`가 포함된 `Email me at camperbot@blog.io` 요소가 있어야 합니다.

```js
const pElement = document.querySelector("footer p");
assert.isNotNull(pElement);
assert.strictEqual(pElement?.textContent, "Email me at camperbot@blog.io");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Camperbot's Blog</title>
</head>
<body>
  <h1>Welcome to Camperbot's Blog</h1>

  <nav>
    <h2>Navigation</h2>
    <ul>
      <li><a href="#post1">My Journey</a></li>
      <li><a href="#post2">Accessibility</a></li>
      <li><a href="#post3">Next Steps</a></li>
    </ul>
  </nav>

  <main>
    <article>
      <h2 id="post1">My Journey Learning to Code</h2>
      <p>I started learning to code a few months ago and it's been a wild ride!</p>
      
      <h3>Early Challenges</h3>
      <p>At first, syntax was really confusing.</p>
      
      <h3>Breakthroughs</h3>
      <p>Eventually things started to click.</p>
    </article>

    <article>
      <h2 id="post2">Accessibility Matters</h2>
      <p>Today I learned that not everyone uses the web the same way I do.</p>

      <h3>Screen Readers</h3>
      <p>These tools help visually impaired users browse websites.</p>
    </article>

    <article>
      <h2 id="post3">What's Next?</h2>
      <p>I'm excited to dive into JavaScript and build interactive features!</p>

      <h3>Coming soon: My first JavaScript project!</h3>
      <p>Stay tuned for some exciting interactive blog features.</p>
    </article>
  </main>
--fcc-editable-region--
  <h3>Contact Me</h3>
  <p>Email me at camperbot@blog.io</p>
--fcc-editable-region--
</body>
</html>
```
