---
id: 68c318d9c3b3f382f93d1fd1
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

페이지의 주요 제목 아래에는 각 블로그 게시물로 연결되는 링크가 포함된 내비게이션 섹션이 있습니다.

이 섹션은 내비게이션 랜드마크 역할을 나타내기 위해 `nav` 요소로 감싸야 합니다.

# --hints--

페이지에 `nav` 요소가 있어야 합니다.

```js
const navElement = document.querySelector("nav");
assert.isNotNull(navElement);
```

`<h2>Navigation</h2>`는 `nav` 요소 안에 있어야 합니다.

```js
const h2Element = document.querySelector("nav h2");
assert.isNotNull(h2Element);
assert.strictEqual(h2Element?.textContent, "Navigation");
```

`ul` 요소는 `nav` 요소 안에 있어야 합니다.

```js
const ulElement = document.querySelector("nav ul");
assert.isNotNull(ulElement);
```

`li` 요소들은 `ul` 요소 안에 있어야 합니다.

```js
const liElements = document.querySelectorAll("nav ul li");
assert.isAtLeast(liElements.length, 3);
```

추가로 `ul`, `li` 또는 `h2` 요소를 추가하면 안 됩니다.

```js
const liElements = document.querySelectorAll("nav ul li");
assert.lengthOf(liElements, 3);

const ulElements = document.querySelectorAll("ul");
assert.lengthOf(ulElements, 1);

const h2Elements = document.querySelectorAll("h2");
assert.lengthOf(h2Elements , 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Camperbot's Blog</title>
</head>
<body>
  <h1>Welcome to Camperbot's Blog</h1>
--fcc-editable-region--
  <h2>Navigation</h2>
  <ul>
    <li><a href="#post1">My Journey</a></li>
    <li><a href="#post2">Accessibility</a></li>
    <li><a href="#post3">Next Steps</a></li>
  </ul>
--fcc-editable-region--
  <div>
    <h2 id="post1">My Journey Learning to Code</h2>
    <p>I started learning to code a few months ago and it's been a wild ride!</p>
    
    <h3>Early Challenges</h3>
    <p>At first, syntax was really confusing.</p>
    
    <h3>Breakthroughs</h3>
    <p>Eventually things started to click.</p>
  </div>

  <div>
    <h2 id="post2">Accessibility Matters</h2>
    <p>Today I learned that not everyone uses the web the same way I do.</p>

    <h3>Screen Readers</h3>
    <p>These tools help visually impaired users browse websites.</p>
  </div>

  <div>
    <h2 id="post3">What's Next?</h2>
    <p>I'm excited to dive into JavaScript and build interactive features!</p>

    <h3>Coming soon: My first JavaScript project!</h3>
    <p>Stay tuned for some exciting interactive blog features.</p>
  </div>
  <h3>Contact Me</h3>
  <p>Email me at camperbot@blog.io</p>
</body>
</html>
```
