---
id: a3bfc1673c0526e06d3ac698
title: 소수 합계 계산기 만들기
challengeType: 26
dashedName: build-a-prime-number-sum-calculator
---

# --description--

이 실습에서는 숫자를 입력받아 그 숫자 이하의 모든 소수 합계를 반환하는 계산기를 만드세요.

**목표**: 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 숫자를 인수로 받는 `sumPrimes` 함수를 만들어야 합니다.
1. `sumPrimes` 함수는 입력된 숫자 이하의 모든 소수 합계를 반환해야 합니다.
1. 입력 숫자가 `2`보다 작으면 함수는 `0`을 반환해야 합니다.

# --hints--

`sumPrimes` 함수를 만들어야 합니다.

```js
assert.isFunction(sumPrimes);
```

`sumPrimes(10)`는 `17`을 반환해야 합니다.

```js
assert.strictEqual(sumPrimes(10), 17);
```

`sumPrimes(5)`는 `10`을 반환해야 합니다.

```js
assert.strictEqual(sumPrimes(5), 10);
```

`sumPrimes(2)`는 `2`을 반환해야 합니다.

```js
assert.strictEqual(sumPrimes(2), 2);
```

`sumPrimes(0)`는 `0`을 반환해야 합니다.

```js
assert.strictEqual(sumPrimes(0), 0);
```

`sumPrimes(977)`는 `73156`을 반환해야 합니다.

```js
assert.strictEqual(sumPrimes(977), 73156);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumPrimes(num) {
  function isPrime(n) {
    if (n < 2) return false;
    for (let i = 2; i <= Math.sqrt(n); i++) {
      if (n % i === 0) {
        return false;
      }
    }
    return true;
  }

  let sum = 0;

  for (let i = 2; i <= num; i++) {
    if (isPrime(i)) {
      sum += i;
    }
  }
  return sum;
}
```
