---
id: 669fcb06c3034a39f5431a38
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

To prepare creating some actual content, add a `body` element below the `head` element.

# --hints--

`<body>` の開始タグが 1 つ必要です。

```js
assert.match(code, /<body>/i);
```

終了タグ `</body>` が 1 つ必要です。

```js
assert.match(code, /<\/body>/i);
```

`head` 要素を変更しないでください。 終了タグを削除していないか確認してください。

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

`body` 要素を `head` 要素の後に置く必要があります。

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>

  --fcc-editable-region--
</html>
```
