---
id: 669fc7e141e4703748c558bf
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

Inside the `html` element, add a `head` element.

# --hints--

`head` の開始タグが 1 つ必要です。

```js
assert.match(code, /<head>/i);
```

`head` の終了タグが 1 つ必要です。

```js
assert.match(code, /<\/head>/i);
```

Your opening `head` tag should come before the closing `head` tag.

```js
assert.match(code, /<head>[.\n\s]*<\/head>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
