---
id: 69899123889dcf5b19be1847
title: Implementa el Algoritmo de Búsqueda en Profundidad
challengeType: 26
dashedName: implement-depth-first-search-algorithm
---

# --description--

En este laboratorio, implementarás un algoritmo de recorrido de grafos llamado <dfn>búsqueda en profundidad</dfn> (DFS).

Mientras que una <dfn>búsqueda en anchura</dfn> (BFS) explora los vecinos nivel por nivel, una <dfn>búsqueda en profundidad</dfn> se adentra lo más profundo posible por un solo camino de aristas antes de retroceder.

Una vez que el algoritmo llega al final de un camino (un nodo sin vecinos no visitados), retrocede al nodo más reciente que aún tiene aristas no exploradas y continúa la búsqueda desde allí.

Para implementar este algoritmo de forma iterativa, querrás usar una **pila**. En JavaScript, un `Array` estándar funciona perfectamente como una pila usando los métodos `.push()` y `.pop()`, siguiendo el principio **Last-In-First-Out (LIFO)**. Esto asegura que el vecino agregado más recientemente a la pila sea el siguiente en ser explorado.

Un salida simple de este algoritmo es una lista de nodos los cuales son alcanzables desde un nodo determinado. Por lo tanto, también querrá llevar un seguimiento de los nodos que visita.

**Objetivo**: Cumple las historias de usuario a continuación y logra que todas las pruebas pasen para completar el laboratorio.

**Historias de usuario:**

1. Debes definir una función llamada `dfs`.
1. La función `dfs` debe tomar dos argumentos:

    - `graph`: Una matriz de adyacencia (un arreglo de arreglos).
    - `root`: Una etiqueta numérica de nodo (el índice inicial) que es el valor numérico del nodo entre `0` y `n - 1`, donde `n` es el número total de nodos en el grafo.

1. La función debe devolver un arreglo que contenga todas las etiquetas de nodo accesibles desde el nodo inicial.

# --hints--

Debes tener una función llamada `dfs` que tome dos argumentos.

```js
assert.isFunction(dfs);
assert.lengthOf(dfs, 2);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)` debería devolver una lista con `1`, `2`, `3` y `0`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)` debería devolver una lista con `1`, `2`, `3` y `0`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)` debería devolver `[3]`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3), [3]);
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)` debería devolver una lista con `3` y `2`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3), [3, 2]);

```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)` debería devolver una lista con `0` y `1`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0), [0, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfs(graph, root) {
  let stack = [root];
  let visited = [];

  while (stack.length > 0) {
    let current = stack.pop();

    if (!visited.includes(current)) {
      visited.push(current);
      let neighbors = graph[current];
      
      for (let i = 0; i < neighbors.length; i++) {
        if (neighbors[i] === 1) {
          stack.push(i);
        }
      }
    }
  }
  return visited;
}
```
