---
id: a202eed8fc186c8434cb6d61
title: 文字列を反転させる関数を作ろう
challengeType: 26
dashedName: build-a-string-inverter
---

# --description--

このラボでは、与えられた文字列の文字を逆順にするシンプルな文字列反転関数を作成します。

例えば、`"hello"` は `"olleh"` になります。

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー:**

1. `reverseString` という名前の関数を作成し、文字列を引数として受け取る必要があります。
2. その関数は、逆順にした文字列を返す必要があります。

# --hints--

`reverseString` という名前の関数を用意してください。

```js
assert.isFunction(reverseString);
```

`reverseString` は文字列を引数として受け取る必要があります。

```js
assert.match(reverseString.toString(), /\s*function\s+reverseString\s*\(\s*\w+\s*\)/);
```

`reverseString("hello")` は文字列を返す必要があります。

```js
assert.isString(reverseString('hello'));
```

`reverseString("hello")` は文字列 `olleh` を返す必要があります。

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

`reverseString("Howdy")` は文字列 `ydwoH` を返す必要があります。

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

`reverseString("Greetings from Earth")` は文字列 `htraE morf sgniteerG` を返す必要があります。

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
