---
id: 6720b3191bbafd269bae0ae8
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

أنشئ عنصر `div` له `id` باسم `greeting-card` و`class` باسم `card`.

داخل عنصر الـ `div`، أضف عنصر `h1` يحتوي على النص `Happy Birthday!`. ثم أضف عنصر فقرة له `class` باسم `message` والنص `Wishing you all the happiness and joy on your special day!`.

# --hints--

يجب أن يكون لديك عنصر `div`.

```js
assert.exists(document.querySelector("div"));
```

يجب أن يحتوي عنصر الـ `div` على `id` باسم `greeting-card`.

```js
assert.strictEqual(document.querySelector("div").id, "greeting-card");
```

يجب أن يحتوي عنصر الـ `div` على `class` باسم `card`.

```js
assert.strictEqual(document.querySelector("#greeting-card").className, "card");
```

يجب أن يكون لديك عنصر `h1` داخل عنصر الـ `div`.

```js
assert.exists(document.querySelector("div > h1"));
```

يجب أن يحتوي عنصر الـ `h1` على النص `Happy Birthday!`.

```js
assert.strictEqual(document.querySelector("div > h1").textContent.trim(), "Happy Birthday!");
```

يجب أن يكون لديك عنصر `p` داخل الـ `div`.

```js
assert.exists(document.querySelector("div > p"));
```

يجب أن يحتوي عنصر الـ `p` على `class` باسم `message`.

```js
assert.exists(document.querySelector("div > p.message"));
```

يجب أن يحتوي عنصر الـ `p` على النص `Wishing you all the happiness and joy on your special day!`.

```js
assert.strictEqual(document.querySelector("div > p.message").textContent.trim(), "Wishing you all the happiness and joy on your special day!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css

```
