---
id: 6720b7f3155a9278d6b6a8b8
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

حان الوقت الآن لتنسيق بطاقة التحية الخاصة بك.

أضف محدِّدًا لعنصر `body`، ثم:

- غيّر `font-family` إلى `Arial` متبوعة بـ `sans-serif` العامة،
- أعطِ حشوًا بمقدار `40px` من جميع الجهات،
- اضبط خاصية `text-align` على `center`.

# --hints--

يجب أن يكون لديك محدِّد `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

يجب أن تضبط `font-family` إلى `"Arial", sans-serif` في محدِّد `body`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("font-family"), ["'Arial', sans-serif", "Arial, sans-serif", '"Arial", sans-serif']);
```

يجب أن تضبط `padding` إلى `40px` في محدِّد `body`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("padding"), "40px");
```

يجب أن تضبط `text-align` إلى `center` في محدِّد `body`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("text-align"), "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
