---
id: 672b4434dd6d6a1edc405419
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

أضف عنصرين `a` داخل العنصر `.card-links`.

يجب أن يحتوي الأول على نص `Send Card`، و`class` بقيمة `send-link`، و`href` بقيمة `#send`.

يجب أن يحتوي الثاني على نص `Share on Social Media`، و`class` بقيمة `share-link`، و`href` بقيمة `#share`.

# --hints--

يجب أن يحتوي أول عنصر `a` داخل `.card-links` على نص `Send Card`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].innerText, "Send Card");
```

يجب أن يحتوي أول عنصر `a` على `class` بقيمة `send-link`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].className, "send-link");
```

يجب أن يحتوي أول عنصر `a` على `href` بقيمة `#send`.

```js
assert.match(document.querySelectorAll(".card-links > a")[0].href, /#send$/);
```

يجب أن يحتوي ثاني عنصر `a` داخل `.card-links` على نص `Share on Social Media`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].innerText, "Share on Social Media");
```

يجب أن يحتوي ثاني عنصر `a` على `class` بقيمة `share-link`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].className, "share-link");
```

يجب أن يحتوي ثاني عنصر `a` على `href` بقيمة `#share`.

```js
assert.match(document.querySelectorAll(".card-links > a")[1].href, /#share$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
