---
id: 672e44e811e610232fead333
title: خطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

أضف عنصر `h2` إلى `#send` يحتوي على النص `Sending your card...`، ثم أضف عنصر `p` يحتوي على النص `Card successfully sent to your recipient!`.

# --hints--

يجب أن يحتوي عنصر `#send` على عنصر `h2`.

```js
assert.exists(document.querySelector("#send > h2"));
```

يجب أن يحتوي عنصر `h2` على النص `Sending your card...`.

```js
assert.strictEqual(document.querySelector("#send > h2").innerText.trim(), "Sending your card...");
```

العنصر الثاني داخل `#send` يجب أن يكون عنصر `p`.

```js
assert.strictEqual(document.querySelector("#send").children[1].tagName, "P");
```

يجب أن يحتوي عنصر `p` على النص `Card successfully sent to your recipient!`.

```js
assert.strictEqual(document.querySelector("#send").children[1].innerText.trim(), "Card successfully sent to your recipient!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
  --fcc-editable-region--
      
  --fcc-editable-region--
    </section>
    <section id="share"></section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
