---
id: 66ebd4ae2812430bb883c787
title: بناء مركز أحداث
challengeType: 25
dashedName: lab-event-hub
demoType: onClick
---

# --description--

في هذا المختبر ستستخدم عناصر HTML الدلالية لإنشاء هيكل صفحة ويب. ستضيف محتوى وصورًا لجعلها تبدو كمركز أحداث حقيقي.

**الهدف:** أتمم قصص المستخدم أدناه واجتز جميع الاختبارات لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك عنصر `header`.
1. داخل عنصر `header`، يجب أن يكون لديك عنصر `h1` يحتوي على النص `Event Hub`، وعنصر `nav`.
1. داخل عنصر `nav`، يجب أن يكون لديك قائمة غير مرتبة من عنصرين يحتويان على روابط لأقسام مختلفة من الصفحة. يجب أن يحتوي العنصر الأول على النص `Upcoming Events`، والعنصر الثاني على النص `Past Events`.
1. يجب تمثيل كل رابط بواسطة عنصر `a` مع خاصية `href` تربط بالقسم المقابل من الصفحة، `#upcoming-events` و `#past-events` على التوالي.
1. يجب أن يكون لديك عنصر `main` يحتوي على الأقسام المختلفة للصفحة.
1. داخل عنصر `main`، يجب أن يكون لديك عنصران `section`.
1. يجب أن يحتوي العنصر الأول `section` على خاصية `id` بقيمة `upcoming-events`
1. داخل قسم `#upcoming-events`، يجب أن يكون لديك:
   
- عنصر `h2` يحتوي على النص `Upcoming Events`.
   - عنصران `article`. يجب أن يمثل كل مقال حدثًا، ويجب أن يحتوي على:
      - عنصر `h3` لعنوان الحدث.
      - عنصر `p` لوصف الحدث. يمكنك إضافة تاريخ في الأسفل إذا أردت.

1. يجب أن يحتوي العنصر الثاني `section` على خاصية `id` بقيمة `past-events`.
1. داخل قسم `#past-events`، يجب أن يكون لديك:

- عنصر `h2` يحتوي على النص `Past Events`.
   - عنصران `article`. يجب أن يمثل كل مقال حدثًا سابقًا، ويجب أن يحتوي على:
     - عنصر `h3` لعنوان الحدث،
     - عنصر `p` لوصف الحدث. يمكنك إضافة تاريخ في الأسفل إذا أردت.
     - عنصر صورة مع خاصية `src` تشير إلى ملف صورة وخاصية `alt` تحتوي على وصف للصورة.

**ملاحظة:** يمكنك استخدام أي نصوص لوصف الأحداث والتواريخ. يمكنك استخدام عناوين URL التالية للصور إذا أردت:

- `https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg`.
- `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

# --hints--

يجب أن يكون لديك عنصر `header`.

```js
assert.isNotNull(document.querySelector("header"));
```

يجب أن يأتي عنصر `header` بعد وسم فتح `body`.

```js
assert.equal(document.querySelector("body")?.firstElementChild?.tagName, "HEADER");
```

داخل عنصر `header`، يجب أن يكون لديك عنصر `h1` يحتوي على النص `Event Hub`.

```js
const h1Element = document.querySelector('header h1');
assert.strictEqual(h1Element?.innerText.trim(), "Event Hub");
```

داخل عنصر `header`، بعد عنصر `h1`، يجب أن يكون لديك عنصر `nav`.

```js
assert.isNotNull(document.querySelector("header>h1+nav"));
```

يجب أن يحتوي عنصر `nav` على قائمة غير مرتبة من عنصرين.

```js
const liElements = document.querySelectorAll('header nav>ul>li');

assert.isNotNull('header nav>ul');
assert.strictEqual(liElements.length, 2);
```

العنصر الأول في القائمة غير المرتبة يجب أن يكون رابطًا.

```js
const firstLink = document.querySelectorAll('header nav ul li a')[0];
assert.exists(firstLink);
```

العنصر الثاني في القائمة غير المرتبة يجب أن يكون رابطًا.

```js
const secondLink = document.querySelectorAll('header nav ul li a')[1];
assert.exists(secondLink);
```

نص العنصر الأول في القائمة غير المرتبة يجب أن يكون `Upcoming Events`.

```js
const firstLink = document.querySelectorAll('header nav>ul>li>a')[0];
assert.strictEqual(firstLink.innerText.trim(), "Upcoming Events");
```

يجب أن تحتوي الخاصية `href` للعنصر الأول في القائمة غير المرتبة على القيمة `#upcoming-events`.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[0];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#upcoming-events");
```

نص العنصر الثاني في القائمة غير المرتبة يجب أن يكون `Past Events`.

```js
const secondLink = document.querySelectorAll('header nav>ul>li>a')[1];
assert.strictEqual(secondLink.innerText.trim(), "Past Events");
```

يجب أن تحتوي الخاصية `href` للعنصر الثاني في القائمة غير المرتبة على القيمة `#past-events`.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[1];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#past-events");
```

يجب أن يكون لديك عنصر `main` بعد وسم إغلاق عنصر `header`.

```js
const mainElement = document.querySelector("body>header+main");
assert.isNotNull(mainElement);
```

داخل عنصر `main`، يجب أن يكون لديك عنصران `section`.

```js
const sectionElements = document.querySelectorAll('body>header+main>section');
assert.strictEqual(sectionElements.length, 2);
```

يجب أن يحتوي العنصر الأول `section` على خاصية `id` بقيمة `upcoming-events`.

```js
const firstSection = document.querySelectorAll('body>header+main>section')[0];
const idAttribute = firstSection?.getAttribute("id");
assert.strictEqual(idAttribute, "upcoming-events");
```

يجب أن يحتوي العنصر الثاني `section` على خاصية `id` بقيمة `past-events`.

```js
const secondSection = document.querySelectorAll('body>header+main>section')[1];
const idAttribute = secondSection?.getAttribute("id");
assert.strictEqual(idAttribute, "past-events");
```

داخل قسم `#upcoming-events`، يجب أن يكون لديك عنصر `h2` يحتوي على النص `Upcoming Events`.

```js
const h2Element = document.querySelector('#upcoming-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Upcoming Events");
```

داخل قسم `#upcoming-events`، يجب أن يكون لديك عنصران `article` أسفل عنصر `h2`.

```js
const articleElements = document.querySelectorAll('#upcoming-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

يجب أن يحتوي كلا العنصرين `article` داخل قسم `#upcoming-events` على عنصر `h3` لعنوان الحدث.

```js
const h3Elements = document.querySelectorAll('#upcoming-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

يجب أن يحتوي كلا العنصرين `article` داخل قسم `#upcoming-events` على عنصر فقرة لوصف الحدث.

```js
const articles = document.querySelectorAll('#upcoming-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

داخل قسم `#past-events`، يجب أن يكون لديك عنصر `h2` يحتوي على النص `Past Events`.

```js
const h2Element = document.querySelector('#past-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Past Events");
```

داخل قسم `#past-events`، يجب أن يكون لديك عنصران `article` أسفل عنصر `h2`.

```js
const articleElements = document.querySelectorAll('#past-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

يجب أن يحتوي كلا العنصرين `article` داخل قسم `#past-events` على عنصر `h3` لعنوان الحدث.

```js
const h3Elements = document.querySelectorAll('#past-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

يجب أن يحتوي كلا العنصرين `article` داخل قسم `#past-events` على عنصر فقرة لوصف الحدث.

```js
const articles = document.querySelectorAll('#past-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

يجب أن يحتوي كلا العنصرين `article` داخل قسم `#past-events` على عنصر صورة.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);
```

يجب أن تحتوي خاصية `src` لكلا العنصرين صورة داخل قسم `#past-events` على ملف صورة.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("src"));
}
```

يجب أن تحتوي خاصية `alt` لكلا العنصرين صورة داخل قسم `#past-events` على وصف للصورة.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("alt"));
}
```

يجب أن يحتوي كل عنصر `h3` على عنوان الحدث.

```js
const eventTitles = document.querySelectorAll('h3');
assert.isNotEmpty(eventTitles);
eventTitles.forEach((eventTitle => assert.isNotEmpty(eventTitle.innerText.trim())));
```

يجب أن يحتوي كل عنصر `p` على وصف الحدث.

```js
const eventDescriptions = document.querySelectorAll('p');
assert.isNotEmpty(eventDescriptions);
eventDescriptions.forEach((eventDescription => assert.isNotEmpty(eventDescription.innerText.trim())));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Event Hub</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Event Hub</title>
</head>
<body>
    <header>
        <h1>Event Hub</h1>
        <nav>
            <ul>
                <li><a href="#upcoming-events">Upcoming Events</a></li>
                <li><a href="#past-events">Past Events</a></li>
            </ul>
        </nav>
    </header>
    <main>
        <section id="upcoming-events">
            <h2>Upcoming Events</h2>
            <article>
                <h3>AI & Machine Learning Conference 2024</h3>
                <p>Join us for a deep dive into the latest advancements in artificial intelligence and machine learning. Industry leaders will share insights and case studies on how AI is transforming various sectors.</p>
                <p>Date: August 10, 2024</p>
            </article>
            <article>
                <h3>Web Development Bootcamp</h3>
                <p>A hands-on workshop designed for developers looking to enhance their skills in modern web technologies including React, Node.js, and GraphQL. Perfect for both beginners and experienced developers.</p>
                <p>Date: September 5, 2024</p>
            </article>
        </section>
        <section id="past-events">
            <h2>Past Events</h2>
            <article>
                <h3>Cybersecurity Summit 2024</h3>
                <p>An event focusing on the latest trends and threats in cybersecurity. Experts discussed strategies for protecting data and ensuring privacy in an increasingly digital world.</p>
                <p>Date: June 15, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg" alt="Image from Cybersecurity Summit 2024">
            </article>
            <article>
                <h3>Blockchain Expo 2024</h3>
                <p>A comprehensive event covering the future of blockchain technology. Topics included decentralized finance (DeFi), smart contracts, and the impact of blockchain on various industries.</p>
                <p>Date: July 20, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg" alt="Image from Blockchain Expo 2024">
            </article>
        </section>
    </main>
</body>
</html>
```

