---
id: 61696ef7ac756c829f9e4048
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

이 프로젝트에서는 외부 CSS 파일을 사용해 페이지를 스타일링합니다. 이미 `styles.css` 파일을 만들어 두었습니다. 하지만 사용하려면 먼저 페이지에 연결해야 합니다.

`link` 요소 안에 `head` 요소를 중첩하세요. `rel` 속성을 `"stylesheet"`로, `href` 속성을 `"styles.css"`로 설정하세요.

# --hints--

`link` 요소가 하나 있어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

`link` 요소는 빈 요소여야 하며, 닫는 태그 `</link>`가 없어야 합니다.

```js
assert.notMatch(code, /<\/link>/i);
```

`link` 요소에는 값이 `rel`인 `"stylesheet"` 속성이 있어야 합니다.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

`link` 요소에는 값이 `href`인 `"styles.css"` 속성이 있어야 합니다.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

`link` 요소는 `head` 요소 안에 있어야 합니다.

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
