---
id: 616965351e74d4689eb6de30
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

웹 페이지에는 여러 개의 `meta` 요소를 포함할 수 있습니다. 각 `meta` 요소는 다른 HTML 요소로 표현할 수 없는 페이지에 대한 정보를 추가합니다.

`meta` 안에 또 다른 `head` 요소를 추가하세요. `name` 속성을 `"viewport"`로, `content` 속성을 `"width=device-width, initial-scale=1.0"`로 설정하여 모든 기기에서 페이지가 동일하게 보이도록 하세요.

# --hints--

`meta` 요소가 두 개 있어야 합니다.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 2);
```

새로운 `meta` 요소는 빈 요소여야 하며, 닫는 태그 `</meta>`가 없어야 합니다.

```js
assert.notMatch(code, /<\/meta>/i);
```

새로운 `meta` 요소는 `name` 속성을 `"viewport"`로, `content` 속성을 `"width=device-width, initial-scale=1.0"`로 설정해야 합니다.

```js
const meta = [...document.querySelectorAll('meta')];
const contentAttrRegex = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/i;
const target = meta?.find(m => m?.getAttribute('name')?.toLowerCase()?.replace(/\s*/g, '') === 'viewport' && contentAttrRegex.test(m?.getAttribute('content')));
assert.exists(target);
```

두 개의 `meta` 요소는 `head` 요소 안에 있어야 합니다.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('meta').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
