---
id: 616971b3cd990186b66c99a1
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

외부 CSS 파일을 설정했으니 이제 페이지 스타일을 지정할 수 있습니다.

참고로, 단락 요소를 선택하고 오른쪽 정렬하는 방법은 다음과 같습니다.

```css
p {
  text-align: right;
}
```

새로운 CSS 규칙을 만들어 `h1` 요소를 선택하고, `text-align` 속성을 `center`로 설정하세요.

# --hints--

`h1` 선택자를 사용하세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

`h1` CSS 규칙에 `text-align` 속성을 `center`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
