---
id: 61695197ac34f0407e339882
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이전 프로젝트에서 보았듯이, 웹페이지는 `DOCTYPE html` 선언으로 시작하고, 그 다음에 `html` 요소가 와야 합니다.

문서 맨 위에 `DOCTYPE html` 선언을 추가하고, 그 다음에 `html` 요소를 추가하세요. `html` 요소에 `lang` 속성을 주고, 그 값으로 `en`를 설정하세요.

# --hints--

코드에 `DOCTYPE` 참조가 포함되어야 합니다.

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 참조 뒤에는 공백을 포함해야 합니다.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

문서 유형을 `html`로 정의해야 합니다.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

`DOCTYPE` 선언은 타입 뒤에 `>`로 닫아야 합니다.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

`html` 요소는 `lang` 속성이 `en`인 여는 태그를 가져야 합니다.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 요소는 닫는 태그가 있어야 합니다.

```js
assert(code.match(/<\/html\s*>/gi));
```

`DOCTYPE` 선언은 HTML 맨 앞에 있어야 합니다.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
