---
id: 61695c4aad56f95497c19583
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

`title` 요소가 검색 엔진에 페이지에 대한 추가 정보를 제공한다는 점을 기억하세요. 또한 `title` 요소의 내용을 두 가지 방식으로 더 표시합니다:

* 페이지가 열려 있을 때 제목 표시줄에  
* 페이지의 브라우저 탭 위에 마우스를 올리면, 해당 탭이 활성화되어 있지 않아도 `title` 텍스트가 표시됩니다.

`head` 요소 안에 `title` 요소를 중첩하고, 그 안에 `Colored Markers` 텍스트를 넣으세요.

# --hints--

열리는 `title` 태그가 있어야 합니다.

```js
assert(code.match(/<title\s*>/i));
```

닫히는 `title` 태그가 있어야 합니다.

```js
assert(code.match(/<\/title\s*>/i));
```

프로젝트에 `Colored Markers` 제목이 있어야 합니다.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

제목 텍스트의 대소문자와 철자가 정확해야 한다는 점을 기억하세요.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
