---
id: 61695d1fbc003856628bf561
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

`charset` 속성은 문서에서 사용하는 문자 인코딩을 지정합니다. `utf-8` (유니코드 변환 형식 – 8비트)은 전자 통신에 사용되는 문자 인코딩 표준입니다.

`head` 요소 안에 `meta` 요소를 중첩하고, `charset` 속성을 `"utf-8"` 값으로 설정하세요.

# --hints--

`meta` 요소가 하나 있어야 합니다.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 1);
```

`meta` 요소는 빈 요소여야 하며, 닫는 태그 `</meta>`가 없어야 합니다.

```js
assert.notMatch(code, /<\/meta>/i);
```

`meta` 요소는 `charset` 속성을 가지고, 그 값이 `"utf-8"`이어야 합니다.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

`meta` 요소는 `head` 요소 안에 있어야 합니다.

```js
const metaElementRegex = /<head\s*>(?:.|\r|\n)*?<meta(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, metaElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
