---
id: 61695ab9f6ffe951c16d03dd
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`head` 요소 안에 `html` 요소를 중첩하세요. `head` 요소 바로 뒤에 `body` 요소를 추가하세요.

# --hints--

여는 `head` 태그가 있어야 합니다.

```js
assert(code.match(/<head\s*>/i));
```

닫는 `head` 태그가 있어야 합니다.

```js
assert(code.match(/<\/head\s*>/i));
```

여는 `body` 태그가 있어야 합니다.

```js
assert(code.match(/<body\s*>/i));
```

닫는 `body` 태그가 있어야 합니다.

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 여는 태그는 하나만 있어야 합니다.

```js
assert(code.match(/<head\s*>/ig).length === 1);
```

`head` 닫는 태그는 하나만 있어야 합니다.

```js
assert(code.match(/<\/head\s*>/ig).length === 1);
```

`body` 여는 태그는 하나만 있어야 합니다.

```js
assert(code.match(/<body\s*>/ig).length === 1);
```

`body` 닫는 태그는 하나만 있어야 합니다.

```js
assert(code.match(/<\/body\s*>/ig).length === 1);
```

`head` 요소는 비어 있어야 합니다.

```js
assert(code.match(/<head\s*>\s*<\/head\s*>/i));
```

`body` 요소는 비어 있어야 합니다.

```js
assert(code.match(/<body\s*>\s*<\/body\s*>/i));
```

`body` 요소는 `head` 요소 뒤에 위치해야 합니다.

```js
assert(code.match(/<\/head\s*>\s*<body\s*>/i));
```

`head` 요소는 `html` 요소 안에 있어야 합니다.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 요소는 `html` 요소 안에 있어야 합니다.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
