---
id: 616d525007b8c5d8b3308b1c
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

마커가 페이지 중앙에 위치하면 더 보기 좋습니다. 이를 쉽게 하는 방법은 `margin` <dfn>단축 속성</dfn>을 사용하는 것입니다.

지난 프로젝트에서는 `margin-top`와 `margin-left` 같은 속성으로 요소의 여백 영역을 각각 설정했습니다. `margin` 단축 속성은 여러 여백 영역을 한 번에 설정하기 쉽게 만듭니다.

마커를 페이지 중앙에 배치하려면 `margin` 속성을 `auto`로 설정하세요. 그러면 `margin-top`, `margin-right`, `margin-bottom`, `margin-left`가 모두 `auto`로 설정됩니다.

# --hints--

`.marker` CSS 규칙에 `margin` 속성이 `auto`로 설정되어 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
}
--fcc-editable-region--

```
