---
id: 616d50b93ba424d6282c99cf
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

배경색이 적용되었지만, 마커 `div` 요소에 내용이 없어서 기본적으로 높이가 없습니다.

`.marker` CSS 규칙에서 `height` 속성을 `25px`으로, `width` 속성을 `200px`로 설정하세요.

# --hints--

`.marker` CSS 규칙에 `width` 속성을 `200px`으로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.width === '200px');
```

`.marker` CSS 규칙에 `height` 속성을 `25px`으로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.height === '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  background-color: red;
}
--fcc-editable-region--

```
