---
id: 61764f7e5240eb7ccc7f6a0a
title: 18단계
challengeType: 0
dashedName: step-18
---

# --description--

그다음, 클래스 `one`를 대상으로 하는 새로운 CSS 규칙을 만들고 그 `background-color` 속성을 `red`로 설정하세요.

# --hints--

클래스 `one`를 대상으로 하려면 클래스 선택자를 사용하세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one'));
```

`.one` CSS 규칙에 `background-color` 속성을 `red`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--

--fcc-editable-region--

```
