---
id: 617b8b38f32bf2080a140675
title: 22단계
challengeType: 0
dashedName: step-22
---

# --description--

함수는 입력을 받아 특정 작업을 수행하는 코드 조각입니다. CSS `rgb` 함수는 빨강, 초록, 파랑에 대한 값, 즉 <dfn>인수</dfn>를 받아 색상을 만듭니다:

```css
rgb(red, green, blue);
```

각 빨강, 초록, 파랑 값은 `0`부터 `255`까지의 숫자입니다. `0`은 해당 색상이 0%임을 의미하며 검정색입니다. `255`는 해당 색상이 100%임을 의미합니다.

`.one` CSS 규칙에서 색상 키워드 `red`를 `rgb` 함수로 바꾸세요. `rgb` 함수에 대해 빨강 값은 `255`, 초록 값은 `0`, 파랑 값은 `0`로 설정하세요.

# --hints--

`.one` CSS 규칙은 `red` 색상 키워드를 사용해 `background-color` 속성을 설정하면 안 됩니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor !== 'red');
```

`.one` CSS 규칙에는 `background-color` 속성이 `rgb(255, 0, 0)`로 설정되어 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: red;
}
--fcc-editable-region--

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
