---
id: 617bb5624a75e86463b7e638
title: 28단계
challengeType: 0
dashedName: step-28
---

# --description--

<dfn>보조 색상</dfn>은 기본 색상을 조합해서 얻는 색상입니다. 빨강, 초록, 파랑 값을 변경하면서 지난 단계에서 몇 가지 보조 색상을 보셨을 수도 있습니다.

첫 번째 보조 색상인 노란색을 만들려면 `rgb` CSS 규칙 안의 `.one` 함수를 순수한 빨강과 순수한 초록을 조합하도록 업데이트하세요.

# --hints--

`.one` CSS 규칙에 `background-color` 속성이 `rgb(255, 255, 0)`로 설정되어 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.two {
  background-color: rgb(0, 255, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
