---
id: 617b97abd9f3f61d1590b815
title: 26단계
challengeType: 0
dashedName: step-26
---

# --description--

가산 RGB 색상 모델에서 <dfn>기본 색상</dfn>은 결합했을 때 순수한 흰색을 만드는 색상입니다. 하지만 이 효과가 나타나려면 각 색상이 최대 강도여야 합니다.

색상을 결합하기 전에 녹색 마커를 순수한 녹색으로 다시 설정하세요. `rgb` CSS 규칙의 `.two` 함수에서 녹색을 `255`의 최대값으로 다시 설정하세요.

# --hints--

`.two` CSS 규칙에 `background-color` 속성이 `rgb(0, 255, 0)`로 설정되어 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 127, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 255);
}

```
