---
id: 617bc3386dc7d07d6469bf20
title: 32단계
challengeType: 0
dashedName: step-32
---

# --description--

주황색을 만들려면 빨간색의 강도를 높이고 녹색 `rgb` 값을 줄여야 했다는 점을 주목하세요. 이는 주황색이 빨간색과 노란색의 조합이기 때문입니다.

3차 색상인 스프링 그린을 만들려면 시안과 녹색을 결합하세요. `rgb` CSS 규칙의 `.two` 함수를 업데이트하여 녹색이 최대값이 되도록 하고 파란색을 `127`으로 설정하세요.

# --hints--

`.two` CSS 규칙에 `background-color` 속성이 `rgb(0, 255, 127)`으로 설정되어 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 255, 255);
}
--fcc-editable-region--

.three {
  background-color: rgb(255, 0, 255);
}

```
