---
id: 617b8e0d93a8d10d9a90e720
title: 23단계
challengeType: 0
dashedName: step-23
---

# --description--

마커의 `background-color`가 여전히 빨간색인 것을 확인하세요. 이는 `rgb` 함수에서 빨간색 값을 `255`의 최대값, 즉 100% 빨간색으로 설정하고, 초록색과 파란색 값은 `0`으로 설정했기 때문입니다.

이제 `rgb` 함수를 사용해 다른 색상을 설정하세요.

`.two` CSS 규칙에서 `rgb` 함수를 사용해 `background-color`를 초록색 최대값으로 설정하고, 나머지 값은 `0`으로 설정하세요. 그리고 `.three` CSS 규칙에서는 `rgb` 함수를 사용해 `background-color`를 파란색 최대값으로 설정하고, 나머지 값은 `0`으로 설정하세요.

# --hints--

`.two` CSS 규칙에서는 `green` 색상 키워드를 사용해 `background-color` 속성을 설정하지 마세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor !== 'green');
```

`.two` CSS 규칙에는 `background-color` 속성이 `rgb(0, 255, 0)`으로 설정되어 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

`.three` CSS 규칙에서는 `blue` 색상 키워드를 사용해 `background-color` 속성을 설정하지 마세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor !== 'blue');
```

`.three` CSS 규칙에는 `background-color` 속성이 `rgb(0, 0, 255)`으로 설정되어 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
