---
id: 617bbb6b97a83f6d1f7d6e38
title: 31단계
challengeType: 0
dashedName: step-31
---

# --description--

이제 보조 색상에 익숙해졌으니 <dfn>삼차 색상</dfn>을 만드는 방법을 배우세요. 삼차 색상은 기본 색상과 인접한 보조 색상을 결합해 만듭니다.

삼차 색상인 주황색을 만들려면 `rgb` CSS 규칙의 `.one` 함수를 업데이트하여 빨간색을 최대값으로 설정하고 녹색을 `127`로 설정하세요.

# --hints--

`.one` CSS 규칙에 `background-color` 속성이 `rgb(255, 127, 0)`로 설정되어 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 255, 0);
}
--fcc-editable-region--

.two {
  background-color: rgb(0, 255, 255);
}

.three {
  background-color: rgb(255, 0, 255);
}

```
