---
id: 617bc4824e233180553a8069
title: 33단계
challengeType: 0
dashedName: step-33
---

# --description--

3차 색상 보라를 만들려면 마젠타와 파란색을 섞으세요. `rgb` CSS 규칙의 `.three` 함수를 업데이트하여 파란색이 최대값이 되도록 하고, 빨간색은 `127`로 설정하세요.

# --hints--

`.three` CSS 규칙에 `background-color` 속성을 `rgb(127, 0, 255)`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(127, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

--fcc-editable-region--
.three {
  background-color: rgb(255, 0, 255);
}
--fcc-editable-region--

```
