---
id: 618a00ed1ca871a2b3aca0cb
title: 35단계
challengeType: 0
dashedName: step-35
---

# --description--

색상환에서 기본색, 보조색, 제3색을 모두 살펴봤으니, 다른 색채 이론 개념과 그것들이 디자인에 미치는 영향을 이해하기가 더 쉬워질 것입니다.

먼저, CSS 규칙 `.one`, `.two`, `.three`에서 `rgb` 함수의 값을 조정하여 각 요소의 `background-color`가 순수한 검정색으로 설정되도록 하세요. `rgb` 함수는 가산 색 모델을 사용한다는 점을 기억하세요. 이 모델에서는 색이 검정색에서 시작해 빨강, 초록, 파랑 값이 증가함에 따라 변합니다.

# --hints--

`.one` CSS 규칙에는 `background-color` 속성이 `rgb(0, 0, 0)`로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(0, 0, 0)');
```

`.two` CSS 규칙에는 `background-color` 속성이 `rgb(0, 0, 0)`로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

`.three` CSS 규칙에는 `background-color` 속성이 `rgb(0, 0, 0)`로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(127, 255, 0);
}

.two {
  background-color: rgb(0, 127, 255);
}

.three {
  background-color: rgb(255, 0, 127);
}
--fcc-editable-region--

```
