---
id: 618a1275e873dcc803c2d1aa
title: 38단계
challengeType: 0
dashedName: step-38
---

# --description--

다음으로, `.one` CSS 규칙에서 `rgb` 함수를 사용해 `background-color`를 검은색으로 설정하세요. 그리고 `.two` CSS 규칙에서는 `rgb` 함수를 사용해 `background-color`를 빨간색으로 설정하세요.

# --hints--

`.one` CSS 규칙에는 `background-color` 속성이 `rgb(0, 0, 0)`로 설정되어 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(0, 0, 0)');
```

`.two` CSS 규칙에는 `background-color` 속성이 `rgb(255, 0, 0)`로 설정되어 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  background-color: rgb(0, 255, 255);
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 255, 255);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
