---
id: 618a0927005553b74bfa05ae
title: 36단계
challengeType: 0
dashedName: step-36
---

# --description--

색상환은 비슷한 색상, 즉 <dfn>색조</dfn>가 서로 가까이 있고, 다른 색상은 더 멀리 떨어져 있는 원입니다. 예를 들어, 순수한 빨간색은 장미색과 주황색 사이에 있습니다.

색상환에서 서로 반대에 위치한 두 색상을 <dfn>보색</dfn>이라고 합니다. 두 보색을 결합하면 회색이 되지만, 나란히 배치하면 강한 시각적 대비를 만들어 더 밝게 보입니다.

`rgb` CSS 규칙의 `.one` 함수에서 빨간색 값을 `255`의 최대값으로 설정하여 순수한 빨간색을 만드세요. `rgb` CSS 규칙의 `.two` 함수에서는 녹색과 파란색 값을 `255`의 최대값으로 설정하여 시안을 만드세요.

# --hints--

`.one` CSS 규칙에 `background-color` 속성을 `rgb(255, 0, 0)`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

`.two` CSS 규칙에 `background-color` 속성을 `rgb(0, 255, 255)`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
