---
id: 619b7424f43ec9215e538afe
title: 44단계
challengeType: 0
dashedName: step-44
---

# --description--

다음으로, 두 번째 마커 `two`에서 클래스 `green`를 `div`로 변경하고, 세 번째 마커 `three`에서 클래스 `blue`를 `div`로 변경하세요.

# --hints--

두 번째 마커 `div`에는 클래스 `two`가 없어야 합니다.

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert(!containerSecondChild?.classList?.contains('two'));
```

두 번째 마커 `div`에는 클래스 `marker`와 `green`가 있어야 합니다.

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert(containerSecondChild?.classList?.contains('marker') && containerSecondChild?.classList?.contains('green'));
```

세 번째 마커 `div`에는 클래스 `three`가 없어야 합니다.

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert(!containerThirdChild?.classList?.contains('three'));
```

세 번째 마커 `div`에는 클래스 `marker`와 `blue`가 있어야 합니다.

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert(containerThirdChild?.classList?.contains('marker') && containerThirdChild?.classList?.contains('blue'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      --fcc-editable-region--
      <div class="marker red">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
      --fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}

```
