---
id: 619b74fa777a2b2473c94f82
title: 45단계
challengeType: 0
dashedName: step-45
---

# --description--

CSS 클래스 선택자 `.two`를 새 `green` 클래스에 맞게 업데이트하세요. 그리고 `.three` 클래스 선택자도 새 `blue` 클래스에 맞게 업데이트하세요.

# --hints--

코드에 더 이상 `.two` 클래스 선택자가 없어야 합니다.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.two'));
```

클래스 `green`를 대상으로 하려면 클래스 선택자를 사용하세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green'));
```

`.green` CSS 규칙에 `background-color` 속성이 `rgb(0, 0, 0)`로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 0, 0)');
```

코드에 더 이상 `.three` 클래스 선택자가 없어야 합니다.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.three'));
```

클래스 `blue`를 대상으로 하려면 클래스 선택자를 사용하세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue'));
```

`.blue` CSS 규칙에 `background-color` 속성이 `rgb(0, 0, 0)`로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
