---
id: 619b7c3c83de403126b69c1e
title: 47단계
challengeType: 0
dashedName: step-47
---

# --description--

이미 0부터 9까지인 10진수, 즉 십진법에 익숙할 수 있습니다. 16진수, 즉 16진법 값은 0부터 9까지, 그리고 A부터 F까지입니다:

```js
0, 1, 2, 3, 4, 5, 6, 7, 8, 9, A, B, C, D, E, F
```

16진수 색상에서 `00`는 해당 색상의 0%이고, `FF`는 100%입니다. 따라서 `#00FF00`는 빨간색 0%, 초록색 100%, 파란색 0%를 의미하며 `rgb(0, 255, 0)`와 같습니다.

16진수 색상의 초록색 값을 `7F`로 설정하여 초록색의 강도를 낮추세요.

# --hints--

`.green` CSS 규칙에 `background-color` 속성을 `#007F00`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.green {
  background-color: #00FF00;
}
--fcc-editable-region--

.blue {
  background-color: rgb(0, 0, 0);
}

```
