---
id: 61a4ada3aabeec822b2975d9
title: 60단계
challengeType: 0
dashedName: step-60
---

# --description--

빨간 마커가 훨씬 더 현실적으로 보입니다. 이제 `linear-gradient` 함수와 16진수 색상을 조합해 녹색 마커도 똑같이 만들어 보세요.

`.green` CSS 규칙에서 `background-color` 속성을 `background`으로 바꾸세요.

# --hints--

`.green` CSS 규칙에는 값이 `background`인 `#007F00` 속성이 있어야 합니다.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.green\{.*?background:#007F00/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background: linear-gradient(180deg, rgb(122, 74, 14) 0%, rgb(245, 62, 113) 50%, rgb(162, 27, 27) 100%);
}

--fcc-editable-region--
.green {
  background-color: #007F00;
}
--fcc-editable-region--

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
