---
id: 61a4a7877da33a73a1c1723e
title: 57단계
challengeType: 0
dashedName: step-57
---

# --description--

색상 중단점이 설정되었으니, `linear-gradient` 함수의 각 색상 인수에 서로 다른 빨간색 음영을 적용하세요. 마커의 위쪽과 아래쪽 가장자리 색상은 더 어두워지고, 가운데 색상은 위에 빛이 있는 것처럼 더 밝아질 것입니다.

현재 순수한 빨간색인 첫 번째 색상 인수에 대해, `rgb` 함수를 업데이트하여 빨간색 값은 `122`, 초록색 값은 `74`, 파란색 값은 `14`로 설정하세요.

# --hints--

`.red` CSS 규칙에 `background` 속성이 `linear-gradient(180deg, rgb(122, 74, 14) 0%, rgb(0, 255, 0) 50%, rgb(0, 0, 255) 100%)`로 설정되어 있어야 합니다.

```js
assert.include(['linear-gradient(rgb(122,74,14)0%,rgb(0,255,0)50%,rgb(0,0,255)100%)', 'rgba(0,0,0,0)linear-gradient(rgb(122,74,14)0%,rgb(0,255,0)50%,rgb(0,0,255)100%)repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(180deg, rgb(255, 0, 0) 0%, rgb(0, 255, 0) 50%, rgb(0, 0, 255) 100%);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
