---
id: 61a5d7ef1cfcf45764df07a2
title: 70단계
challengeType: 0
dashedName: step-70
---

# --description--

마커에 올바른 색상이 적용되었으니, 이제 마커 소매를 만들어 보세요. 빨간 마커부터 시작하세요.

빨간 마커 `div` 요소 안에 새로운 `div` 요소를 만들고 `sleeve` 클래스를 지정하세요.

# --hints--

새로운 `div` 요소는 빨간 마커의 `div` 요소 안에 있어야 합니다.

```js
const redMarkerChildren = [...document.querySelector('.red')?.children];
assert(redMarkerChildren.every(child => child?.localName === 'div') && redMarkerChildren.length === 1);
```

새로운 `div` 요소에는 `sleeve` 클래스를 지정해야 합니다.

```js
const redMarkerChild = [...document.querySelector('.red')?.children][0];
assert(redMarkerChild?.classList?.contains('sleeve'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      --fcc-editable-region--
      <div class="marker red">
      </div>
      --fcc-editable-region--
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
