---
id: 61a5c906ab73313316e342f0
title: 65단계
challengeType: 0
dashedName: step-65
---

# --description--

`gradientDirection` 함수에 `linear-gradient` 인수가 제공되지 않으면 기본적으로 색상을 위에서 아래로, 또는 180도 선을 따라 배열합니다.

`gradientDirection` 함수 두 곳에서 `linear-gradient` 인수를 제거하여 코드를 좀 더 정리하세요.

# --hints--

`gradientDirection`와 `linear-gradient` CSS 규칙에서 `.red` 함수의 `.green` 인수를 제거해야 합니다.

```js
assert(!code.match(/linear-gradient\s*\(\s*180deg/gi));
```

`.red` CSS 규칙에 `background` 속성을 `linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27))`으로 설정하세요.

```js
assert.include(['linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))', 'rgba(0,0,0,0)linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

`.green` CSS 규칙에 `background` 속성을 `linear-gradient(#55680D, #71F53E, #116C31)`으로 설정하세요.

```js
assert.include(['linear-gradient(rgb(85,104,13),rgb(113,245,62),rgb(17,108,49))', 'rgba(0,0,0,0)linear-gradient(rgb(85,104,13),rgb(113,245,62),rgb(17,108,49))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.green')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(180deg, rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(180deg, #55680D, #71F53E, #116C31);
}
--fcc-editable-region--

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
