---
id: 61a5d6701ee08953ca375243
title: 68단계
challengeType: 0
dashedName: step-68
---

# --description--

두 번째 색상 인수로, `hsl` 함수를 사용하고 `223`은 색상, `90%`는 채도, `60%`은 밝기 값으로 전달하세요.

# --hints--

`.blue` CSS 규칙에 `background` 속성을 `linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%))`으로 설정해야 합니다.

```js
assert.include(['linear-gradient(rgb(10,66,72),rgb(61,113,245))', 'rgba(0,0,0,0)linear-gradient(rgb(10,66,72),rgb(61,113,245))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.blue')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

--fcc-editable-region--
.blue {
  background: linear-gradient(hsl(186, 76%, 16%));
}
--fcc-editable-region--

```
