---
id: 61a5bfe091060f1d6a629dd0
title: 64단계
challengeType: 0
dashedName: step-64
---

# --description--

색상 중단점이 없어도 녹색 마커의 색상이 빨간색 마커와 같은 지점에서 전환되는 것을 눈치챘을 수 있습니다. 첫 번째 색상은 시작(0%), 두 번째는 중간(50%), 마지막은 그래디언트 선의 끝(100%)에 있습니다.

`linear-gradient` 함수가 이러한 값을 자동으로 계산해 주며, 기본적으로 그래디언트 선을 따라 색상을 고르게 배치합니다.

`.red` CSS 규칙에서 `linear-gradient` 함수의 세 가지 색상 중단점을 제거해 코드를 좀 더 깔끔하게 만드세요.

# --hints--

`.red` CSS 규칙에 `background` 속성이 `linear-gradient(180deg, rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27))`로 설정되어 있어야 합니다.

```js
assert.include(['linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))', 'rgba(0,0,0,0)linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(180deg, rgb(122, 74, 14) 0%, rgb(245, 62, 113) 50%, rgb(162, 27, 27) 100%);
}
--fcc-editable-region--

.green {
  background: linear-gradient(180deg, #55680D, #71F53E, #116C31);
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
