---
id: 61b0936d9e7fc020b43b1bb8
title: 78단계
challengeType: 0
dashedName: step-78
---

# --description--

소매가 사라진 것처럼 보이지만 걱정하지 마세요 — 여전히 있습니다. 새 캡 `div`가 마커 전체 너비를 차지하면서 소매를 다음 줄로 밀어낸 것입니다.

이는 `display` 요소의 기본 `div` 속성이 `block`이기 때문입니다. 그래서 두 `block` 요소가 나란히 있으면 실제 블록처럼 쌓입니다. 예를 들어, 마커 요소들이 모두 위아래로 쌓여 있습니다.

두 `div` 요소를 같은 줄에 배치하려면 `display` 속성을 `inline-block`으로 설정하세요.

`cap`와 `sleeve` 클래스를 모두 선택하는 새 규칙을 만들고 `display`을 `inline-block`로 설정하세요.

# --hints--

`cap`와 `sleeve` 클래스를 모두 선택하려면 클래스 선택자를 사용하세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve') || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap'));
```

CSS 규칙에 `display` 속성을 `inline-block`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve')?.display === 'inline-block' || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap')?.display === 'inline-block');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
}

--fcc-editable-region--

--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
