---
id: 61b093179e7fc020b43b1bb3
title: 73단계
challengeType: 0
dashedName: step-73
---

# --description--

<dfn>불투명도</dfn>는 얼마나 불투명하거나 투명하지 않은지를 나타냅니다. 예를 들어, 단단한 벽은 불투명해서 빛이 통과하지 못합니다. 하지만 유리잔은 훨씬 더 투명해서 유리를 통해 반대편이 보입니다.

CSS의 `opacity` 속성으로 요소가 얼마나 불투명하거나 투명한지 조절할 수 있습니다. 값이 `0` 또는 0%일 때 요소는 완전히 투명해지고, `1.0` 또는 100%일 때는 기본적으로 그렇듯 완전히 불투명해집니다.

`.sleeve` CSS 규칙에서 `opacity` 속성을 `0.5`로 설정하세요.

# --hints--

`.sleeve` CSS 규칙에 `opacity` 속성이 `0.5`로 설정되어 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sleeve')?.opacity === '0.5');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
