---
id: 61b095c79e7fc020b43b1bbb
title: 62단계
challengeType: 0
dashedName: step-62
---

# --description--

두 번째 색상 인수로, 빨강에는 `71`, 초록에는 `F5`, 파랑에는 `3E` 값을 사용하는 16진수 색상 코드를 사용하세요.

# --hints--

`.green` CSS 규칙에 `background` 속성을 `linear-gradient(180deg, #55680D, #71F53E)`로 설정해야 합니다.

```js
assert.include(['linear-gradient(rgb(85,104,13),rgb(113,245,62))', 'rgba(0,0,0,0)linear-gradient(rgb(85,104,13),rgb(113,245,62))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.green')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background: linear-gradient(180deg, rgb(122, 74, 14) 0%, rgb(245, 62, 113) 50%, rgb(162, 27, 27) 100%);
}

--fcc-editable-region--
.green {
  background: linear-gradient(180deg, #55680D);
}
--fcc-editable-region--

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
