---
id: 61b093219e7fc020b43b1bb4
title: 74단계
challengeType: 0
dashedName: step-74
---

# --description--

요소의 불투명도를 설정하는 또 다른 방법은 <dfn>알파 채널</dfn>을 사용하는 것입니다. `opacity` 속성과 비슷하게, 알파 채널은 색상의 투명도 또는 불투명도를 제어합니다.

이미 `opacity` 속성과 이름이 지정된 색상으로 소매의 불투명도를 설정했지만, 다른 CSS 색상 속성에도 알파 채널을 추가할 수 있습니다.

`.sleeve` 규칙 안에서 `opacity` 속성과 값을 제거하세요.

# --hints--

`.sleeve` CSS 규칙에는 `opacity` 속성과 값이 없어야 합니다.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.sleeve')?.opacity);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
  opacity: 0.5;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
