---
id: 61b30995968123ceb6b76167
title: 88단계
challengeType: 0
dashedName: step-88
---

# --description--

그림자의 가장자리가 뚜렷한 것을 확인하세요. 이는 `blurRadius` 속성에 선택적인 `box-shadow` 값이 있기 때문입니다.

```css
box-shadow: offsetX offsetY blurRadius color;
```

만약 `blurRadius` 값이 포함되지 않으면 기본값은 `0`이 되어 가장자리가 뚜렷하게 나타납니다. `blurRadius` 값이 클수록 흐림 효과가 더 강해집니다.

`.green` CSS 규칙에 `box-shadow` 속성을 추가하고, `5px`에는 `offsetX`, `5px`에는 `offsetY`, `5px`에는 `blurRadius`, `green`에는 `color` 값을 지정하세요.

# --hints--

`.green` CSS 규칙에 `box-shadow` 단축 속성이 있어야 하며, 값은 `5px 5px 5px green`이어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.boxShadow === 'green 5px 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
  box-shadow: -5px -5px red;
}

--fcc-editable-region--
.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}
--fcc-editable-region--

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
