---
id: 61b30a286c228bd0c493c09a
title: 89단계
challengeType: 0
dashedName: step-89
---

# --description--

그런데 그림자를 더 넓게 확장하고 싶다면 어떻게 해야 할까요? 선택 사항인 `spreadRadius` 값을 사용하면 가능합니다.

```css
box-shadow: offsetX offsetY blurRadius spreadRadius color;
```

`blurRadius`처럼 `spreadRadius`도 포함하지 않으면 기본값이 `0`입니다.

파란 마커 주위에 5픽셀 그림자를 직접 추가하는 연습을 해보세요.

`.blue` CSS 규칙에 `box-shadow` 속성을 추가하고, `0`에는 `offsetX`, `0`에는 `offsetY`, `0`에는 `blurRadius`, `5px`에는 `spreadRadius`, `blue`에는 `color` 값을 넣으세요.

# --hints--

`.blue` CSS 규칙에는 `box-shadow` 단축 속성이 있어야 하며, 값은 `0 0 0 5px blue`이어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.boxShadow === 'blue 0px 0px 0px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
  box-shadow: -5px -5px red;
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
  box-shadow: 5px 5px 5px green;
}

--fcc-editable-region--
.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}
--fcc-editable-region--

```
