---
id: 5688e62ea601b2482ff8422b
title: Profilabfrage erstellen
challengeType:  26
dashedName: lab-profile-lookup
---

# --description--

In diesem Lab erstellen Sie eine Profilabfrage, die Informationen über Personen in einer Kontaktliste abruft.

Stellen Sie sich für dieses Beispiel vor, es gibt einen Kontakt namens Akira Laine, die `lookUpProfile("Akira", "lastName")` sollte `Laine` zurückgeben.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Lab abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion namens `lookUpProfile` erstellen, die einen Namen und eine Eigenschaft als Argumente entgegennimmt.  
2. Sie sollten Kontaktinformationen aus dem bereitgestellten `contacts`-Array abrufen.  
3. Wenn die Funktion einen Kontaktnamen erhält und die Eigenschaft beim zugehörigen Kontakt existiert, sollte der Wert der Eigenschaft zurückgegeben werden.  
4. Wenn der an die Funktion übergebene Name keinem Kontakt im Kontakte-Array entspricht, sollte die Funktion `"No such contact"` zurückgeben.  
5. Wenn die Eigenschaft bei einem gefundenen Kontakt nicht existiert, sollte die Funktion `"No such property"` zurückgeben.  

# --before-each--

```js
const _contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];

contacts = structuredClone(_contacts); 
```

# --hints--

Das `contacts`-Array sollte weiterhin im globalen Gültigkeitsbereich vorhanden sein. Setzen Sie die Lektion zurück, um es wiederherzustellen, falls Sie es gelöscht haben.

```js
assert.isArray(contacts);
```

Sie sollten eine Funktion namens `lookUpProfile` haben.

```js
assert.isFunction(lookUpProfile);
```

`lookUpProfile("Kristian", "lastName")` sollte einen String zurückgeben.

```js
assert.isString(lookUpProfile("Kristian", "lastName"));
```

`lookUpProfile("Kristian", "lastName")` sollte den String `Vos` zurückgeben.

```js
assert.strictEqual(lookUpProfile('Kristian', 'lastName'), 'Vos');
```

`lookUpProfile("Sherlock", "likes")` sollte `["Intriguing Cases", "Violin"]` zurückgeben.

```js
assert.deepEqual(lookUpProfile('Sherlock', 'likes'), [
  'Intriguing Cases',
  'Violin'
]);
```

`lookUpProfile("Harry", "likes")` sollte ein Array zurückgeben.

```js
assert.isArray(lookUpProfile('Harry', 'likes'));
```

`lookUpProfile("Bob", "number")` sollte den String `No such contact` zurückgeben.

```js
assert.strictEqual(lookUpProfile('Bob', 'number'),'No such contact');
```

`lookUpProfile("Bob", "potato")` sollte den String `No such contact` zurückgeben.

```js
assert.strictEqual(lookUpProfile('Bob', 'potato'), 'No such contact');
```

`lookUpProfile("Akira", "address")` sollte den String `No such property` zurückgeben.

```js
assert.strictEqual(lookUpProfile('Akira', 'address'), 'No such property');
```

Das `contacts`-Array sollte nach Ausführung von `lookUpProfile` unverändert bleiben.

```js
lookUpProfile("Sherlock", "likes");
assert.deepEqual(contacts,_contacts);
```

# --seed--

## --seed-contents--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


```

# --solutions--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


function lookUpProfile(name, prop) {
  for (let i in contacts) {
    if (contacts[i].firstName === name) {
      return contacts[i][prop] || "No such property";
    }
  }
  return "No such contact";
}
```
