---
id: 685239f0cad43c7a909cbd98
title: Debuggen einer Seite zur Tiervermittlung
challengeType: 25
dashedName: lab-debug-pet-adoption-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

Sally, die Besitzerin eines Tiervermittlungsstores, hat ihre erste Webseite erstellt, aber es gibt einige Probleme.

Ihre Aufgabe ist es, alle Fehler zu beheben, damit Sally ihre Seite weiterentwickeln kann.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Lab abzuschließen.

**User Stories:**

1. Sally möchte ein Bild von einigen Katzen verwenden, aber es wird nicht korrekt angezeigt. Sie müssen Folgendes im `img`-Element korrigieren:
   - Ersetzen Sie das `href`-Attribut durch das korrekte Attribut für die Bildquelle.
   - Ersetzen Sie das `att`-Attribut durch das korrekte Attribut, das kurzen, beschreibenden Text für Bilder darstellt.
   - Entfernen Sie das schließende `</img>`-Tag, da `img`-Elemente sogenannte Void-Elemente sind und keine schließenden Tags haben.
2. Sally möchte einige Links verwenden, um Benutzer zu den Seiten für Hunde und Katzen zu leiten. Aber die Links funktionieren nicht richtig. Sie müssen Folgendes in den `a`-Elementen korrigieren:
   - Ersetzen Sie beide `src`-Attribute durch die korrekten Attribute, die zur Angabe von URLs verwendet werden.

# --hints--

Ihr `img`-Element sollte ein `src`-Attribut anstelle des `href`-Attributs haben.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("src"));
```

Ihr `img`-Element sollte ein `alt`-Attribut anstelle des nicht existierenden `att`-Attributs haben.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("alt"));
```

Ihr `img`-Element sollte kein schließendes `</img>`-Tag haben.

```js
assert.notMatch(code, /<\/img>/);
```

Ihr `a`-Element mit dem Text `Visit cats page` muss ein `href`-Attribut anstelle eines `src`-Attributs haben.

```js
const anchors = document.querySelectorAll("a");
const catLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit cats page");
assert.isDefined(catLink);
assert.isTrue(catLink.hasAttribute("href"));
```

Ihr `a`-Element mit dem Text `Visit dogs page` muss ein `href`-Attribut anstelle eines `src`-Attributs haben.

```js
const anchors = document.querySelectorAll("a");
const dogLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit dogs page");
assert.isDefined(dogLink);
assert.isTrue(dogLink.hasAttribute("href"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" att="Two tabby kittens sleeping together on a couch."></img>

<h2>Adopt a cat!</h2>
<a src="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a src="/dogs">Visit dogs page</a>
```

# --solutions--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">

<h2>Adopt a cat!</h2>
<a href="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a href="/dogs">Visit dogs page</a>
```
