---
id: 68ed661e9f463c155bb677a1
title: خطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة ستتدرب على استخدام CSS flexbox من خلال تصميم مجموعة من الصناديق الملونة.

لبدء العمل، داخل عنصر `body` أضف عنصر `header`. قم بتسكين عنصر `h1` داخل عنصر `header` مع النص `Colored Boxes Layout`.

# --hints--

يجب أن يحتوي عنصر `header` على داخل عنصر `body`.

```js
const headerEl = document.querySelector('body > header');
assert.exists(headerEl);
```

يجب أن يحتوي عنصر `header` على عنصر `h1` بداخله.

```js
const h1El = document.querySelector('body > header > h1');
assert.exists(h1El);
```

يجب أن يحتوي عنصر `h1` على النص `Colored Boxes Layout` بداخله.

```js
const h1El = document.querySelector('body > header > h1');
assert.equal(h1El?.textContent.trim(), 'Colored Boxes Layout');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
</head>
<body>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```
