---
id: 68ef0b3237fef991c92d9c19
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

أضف الآن عنصر `div` مع الصنف `flex-container` أسفل الـ `header`. ثم أنشئ ستة عناصر `div` مع الصنف `box` داخل العنصر `div` الذي أنشأته للتو.

تذكّر أن خاصية الصنف تُستخدم غالبًا للإشارة إلى اسم صنف في ورقة الأنماط. في هذه الحالة، ستُنسّق كل عناصر الـ `div` بشكل متساوٍ وفقًا لتعريف الأنماط `.flex-container` و `.box`.

# --hints--

يجب أن يكون لديك عنصر `div` أسفل عنصر الـ `header`.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.exists(divELBelowHeader);
```

يجب أن يحتوي عنصر الـ `div` أسفل عنصر الـ `header` على الصنف `flex-container`.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.equal(divELBelowHeader.className, 'flex-container');
```

يجب أن يكون لديك ستة عناصر `div` مسكنة داخل عنصر `.flex-container`.

```js
const divElsInsideContainer = document.querySelectorAll('div.flex-container > div');
assert.lengthOf(divElsInsideContainer, 6);
```

يجب أن تحتوي كل عناصر الـ `div` الستة داخل عنصر `.flex-container` على الصنف `box`.

```js
const divElsInsideContainerWithBox = document.querySelectorAll('div.flex-container > div.box');
assert.lengthOf(divElsInsideContainerWithBox, 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Colored Boxes</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <header>
    <h1>Colored Boxes Layout</h1>
  </header>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

```
