---
id: 68f46b758eb8ea7043abd73c
title: الخطوة 33
challengeType: 0
dashedName: step-33
---

# --description--

الآن ستنظم الصناديق داخل حاوية flexbox الخاصة بك. أولًا، أنشئ محدِّد `.box1` وأضف له الخاصية `background` بالقيمة `#f16e79`.

بعدها، أضف الخاصية `order` بالقيمة `1`. سترى أن الصندوق الأول ينتقل إلى النهاية. يحدث هذا لأن صنف `box` لديه بالفعل قيمة `order` تساوي `0`، والعناصر ذات القيم الأعلى تظهر لاحقًا.

أضف أيضًا الخاصية `flex-grow` بالقيمة `1`. هذا سيجعل الصندوق ينمو ليملأ المساحة الإضافية في سطره.

# --hints--

يجب أن يكون لديك محدِّد `.box1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box1'));
```

يجب أن يحتوي محدِّد `.box1` على الخاصية `background` بالقيمة `#f16e79`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('background'), 'rgb(241, 110, 121)');
```

يجب أن يحتوي محدِّد `.box1` على الخاصية `order` بالقيمة `1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('order'), '1');
```

يجب أن يحتوي محدِّد `.box1` على الخاصية `flex-grow` بالقيمة `1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('flex-grow'), '1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

--fcc-editable-region--

--fcc-editable-region--
```
