---
id: 68f46e53fdfdc18a773282ba
title: الخطوة 35
challengeType: 0
dashedName: step-35
---

# --description--

داخل ملف `styles.css`، أنشئ محدِّد `.box2` وأضف خاصية `background` بالقيمة `#f4a261` إليه.

أيضًا، أضف خاصية `order` بقيمة `0`.

# --hints--

يجب أن يكون لديك محدِّد `.box2`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box2'));
```

يجب أن يحتوي محدِّد `.box2` على الخاصية `background` بالقيمة `#f4a261`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box2')?.getPropVal('background'), 'rgb(244, 162, 97)');
```

يجب أن يحتوي محدِّد `.box2` على الخاصية `order` بالقيمة `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box2')?.getPropVal('order'), '0');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Colored Boxes</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <header>
    <h1>Colored Boxes Layout</h1>
  </header>
  <div class="flex-container">
    <div class="box box1">
      <h2>Box 1</h2>
      <p>Red</p>
    </div>
    <div class="box box2">
      <h2>Box 2</h2>
      <p>Orange</p>
    </div>
    <div class="box">
      <h2>Box 3</h2>
      <p>Yellow</p>
    </div>
    <div class="box">
      <h2>Box 4</h2>
      <p>Green</p>
    </div>
    <div class="box">
      <h2>Box 5</h2>
      <p>Blue</p>
    </div>
    <div class="box">
      <h2>Box 6</h2>
      <p>Indigo</p>
    </div>
  </div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

.box1 {
  background: #f16e79;
  order: 1; 
  flex-grow: 1;
}

--fcc-editable-region--

--fcc-editable-region--

```
