---
id: 68f91df5776e693049fb8942
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

حان الوقت الآن لإضافة بعض قواعد CSS إلى ملف `styles.css`. ابدأ بإنشاء محدِّد لعنصر `h1`.

وسّط محتوى عنصر `h1` بتعيين خاصية `text-align` إلى `center`. أنشئ أيضًا خاصية `margin-bottom` بالقيمة `10px` لإضافة هامش بين أسفل عنصر `h1` وأي عنصر HTML يأتي تحته.

# --hints--

يجب أن يكون لديك محدِّد `h1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

يجب أن يعيّن محدِّد `h1` الخاصية `text-align` إلى `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('text-align'), 'center');
```

يجب أن يعيّن محدِّد `h1` الخاصية `margin-bottom` إلى `10px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('margin-bottom'), '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Colored Boxes</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <header>
    <h1>Colored Boxes Layout</h1>
  </header>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```

