---
id: 692f314d449329ba8bcd5d2e
title: الخطوة 17
challengeType: 0
dashedName: step-17
---

# --description--

خاصية `border-radius` تُدوّر زوايا حدود العنصر. يمكنك تحديد:

- قيمة واحدة تُطبق على جميع الزوايا الأربع.
- قيمتان، الأولى للزاويتين العلوية اليسرى والسفلية اليمنى، والثانية للزاويتين العلوية اليمنى والسفلية اليسرى.
- ثلاث قيم، تمثل الزاوية العلوية اليسرى، والزاويتين العلوية اليمنى والسفلية اليسرى، والزاوية السفلية اليمنى.
- أربع قيم، تحدد انحناء الزوايا باتجاه عقارب الساعة بدءًا من الزاوية العلوية اليسرى.

أضف الآن الخاصية `border-radius` بالقيمة `5px` إلى المحدِّد `.box`.

# --hints--

يجب أن يحتوي المحدِّد `.box` على الخاصية `border-radius` بالقيمة `5px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('border-radius'), '5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
