---
id: 692f22a2ad44ae887a24c8f9
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

داخل كل عنصر `.box`، أضف عنصر `h2` يحتوي على النص `Box` متبوعًا بمسافة ورقم تسلسلي يبدأ من `1`.

# --hints--

يجب أن يحتوي أول `div` من صنف `box` على عنصر `h2` بنص `Box 1`.

```js
const firstH2El = document.querySelector("div.box:nth-of-type(1) > h2");
assert.equal(firstH2El.textContent, "Box 1");
```

يجب أن يحتوي العنصر `div` الثاني بصنف `box` على عنصر `h2` بالنص `Box 2`.

```js
const secondH2El = document.querySelector("div.box:nth-of-type(2) > h2");
assert.equal(secondH2El.textContent, "Box 2");
```

يجب أن يحتوي العنصر `div` الثالث الذي يحمل الصنف `box` على عنصر `h2` بالنص `Box 3`.

```js
const thirdH2El = document.querySelector("div.box:nth-of-type(3) > h2");
assert.equal(thirdH2El.textContent, "Box 3");
```

يجب أن يكون لديك عنصر `h2` يحتوي على النص `Box 4` داخل العنصر الرابع `div` الذي يحمل الصنف `box`.

```js
const fourthH2El = document.querySelector("div.box:nth-of-type(4) > h2");
assert.equal(fourthH2El.textContent, "Box 4");
```

يجب أن يحتوي العنصر `div` الخامس الذي يحمل الصنف `box` على عنصر `h2` بالنص `Box 5`.

```js
const fifthH2El = document.querySelector("div.box:nth-of-type(5) > h2");
assert.equal(fifthH2El.textContent, "Box 5");
```

يجب أن يحتوي آخر `div` من صنف `box` على عنصر `h2` بالنص `Box 6`.

```js
const sixthH2El = document.querySelector("div.box:nth-of-type(6) > h2");
assert.equal(sixthH2El.textContent, "Box 6");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
		<div class="box">
			
		</div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
