---
id: 692f22933ba1cb86cd75a70c
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

حان الوقت الآن لإنشاء اسم اللون للصناديق. للقيام بذلك، أضف عنصر `p` أسفل كل عنصر `h2` مع أسماء الألوان هذه بالترتيب: `Red`، `Orange`، `Yellow`، `Green`، `Blue`، و`Indigo`. لون واحد لكل صندوق.

# --hints--

يجب إضافة عنصر `p` يحتوي على النص `Red` أسفل عنصر `h2` داخل أول عنصر `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[0];
assert.equal(heading.textContent, "Red");
```

يجب إضافة عنصر `p` يحتوي على النص `Orange` أسفل عنصر `h2` داخل ثاني عنصر `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[1];
assert.equal(heading.textContent, "Orange");
```

يجب إضافة عنصر `p` يحتوي على النص `Yellow` أسفل عنصر `h2` داخل ثالث عنصر `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[2];
assert.equal(heading.textContent, "Yellow");
```

يجب إضافة عنصر `p` يحتوي على النص `Green` أسفل عنصر `h2` داخل رابع عنصر `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[3];
assert.equal(heading.textContent, "Green");
```

يجب إضافة عنصر `p` يحتوي على النص `Blue` أسفل عنصر `h2` داخل خامس عنصر `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[4];
assert.equal(heading.textContent, "Blue");
```

يجب إضافة عنصر `p` يحتوي على النص `Indigo` أسفل عنصر `h2` داخل سادس عنصر `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[5];
assert.equal(heading.textContent, "Indigo");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
    <div class="box">
      <h2>Box 1</h2>
			
    </div>
    <div class="box">
      <h2>Box 2</h2>
			
    </div>
    <div class="box">
      <h2>Box 3</h2>
			
    </div>
    <div class="box">
      <h2>Box 4</h2>
			
    </div>
    <div class="box">
      <h2>Box 5</h2>
			
    </div>
    <div class="box">
      <h2>Box 6</h2>
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
