---
id: 686386e76197458b7dfd52eb
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

أنشئ عنصر `div` تحت عنصر `h1`. داخل ذلك العنصر `div`، أنشئ عنصر `div` آخر، وقم بتسكين عنصر `div` آخر بداخله.

وبداخل أعمق عنصر `div`، أنشئ عنصر `h2` يحتوي على النص `Listener`. هذا سيمثل خطة التسعير الأولى.

# --hints--

يجب أن تنشئ عنصر `div`.

```js
assert.exists(document.querySelector("div"));
```

يجب أن تنشئ عنصر `div` آخر داخل العنصر الخارجي `div`.

```js
assert.exists(document.querySelector("div > div"));
```

يجب أن تنشئ عنصر `div` آخر داخل العنصر الداخلي `div`.

```js
assert.exists(document.querySelector("div > div > div"));
```

يجب أن تنشئ عنصر `h2` يحتوي على النص `Listener` داخل أعمق عنصر `div`.

```js
assert.equal(document.querySelector("div > div > div > h2")?.textContent, "Listener");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
