---
id: 686386e76197458b7dfd52ef
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

الآن ستُنشئ زرًا للدعوة إلى اتخاذ إجراء (CTA). الدعوة إلى اتخاذ إجراء (CTA) هي زر أو رابط مصمم لتحفيز المستخدم على تنفيذ إجراء محدد، في هذه الحالة، بدء الاستماع.

خارج أعمق عنصر `div`، أنشئ عنصر وسم الرابط مع خاصية `href` بقيمة `#` والنص `Start listening`.

بهذا نكون قد أنهينا شفرة HTML لبطاقة التسعير الأولى.

# --hints--

يجب أن تنشئ عنصر وسم الرابط.

```js
assert.exists(document.querySelector("a"))
```

يجب أن يحتوي عنصر وسم الرابط على خاصية `href` بقيمة `#`.

```js
const link = document.querySelector("a")
assert.equal(link?.getAttribute("href"), "#")
```

يجب أن يحتوي عنصر وسم الرابط على النص `Start listening`.

```js
assert.equal(document.querySelector("a")?.textContent, "Start listening")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          --fcc-editable-region--
          
          --fcc-editable-region-- 
        </div>
      </div>
    </main>
  </body>
</html>
```
