---
id: 686386e76197458b7dfd52ec
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

أسفل العنصر `h2`، أنشئ عنصر `p` يحتوي على النص `$0`. ثم، داخل عنصر `p` وبعد النص، أنشئ عنصر `span` يحتوي على النص `/month`.

# --hints--

يجب أن تنشئ عنصر `p` يحتوي على النص `$0`.

```js
const p = document.querySelector("p");
const pTextBeforeSpan = p.childNodes[0].textContent.trim();

assert.equal(pTextBeforeSpan, "$0");
```

يجب أن تنشئ عنصر `span` يحتوي على النص `/month` داخل عنصر `p`.

```js
const spanText = document.querySelector("p > span")?.textContent;
assert.equal(spanText, "/month");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            --fcc-editable-region--
            
            --fcc-editable-region--
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
