---
id: 686383abac6c625a84458ea9
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستتعلم المزيد عن Tailwind CSS باستخدام فئات الأدوات الشبكية لبناء مكوِّن التسعير.

تم توفير قالب HTML الأساسي لك. ابدأ بإنشاء عنصر `h1` يحتوي على النص `Choose your listening plan`.

# --hints--

يجب أن تنشئ عنصر `h1`.

```js
assert.exists(document.querySelector("h1"));
```

يجب أن يحتوي عنصر `h1` الخاص بك على النص `Choose your listening plan`.

```js
assert.equal(document.querySelector("h1")?.textContent, "Choose your listening plan");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </main>
  </body>
</html>
```
