---
id: 686386e76197458b7dfd52ee
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

داخل العنصر `ul`، أنشئ عنصرين من عناصر القائمة. داخل عنصر القائمة الأول، أنشئ عنصر `span` وأضف كيان HTML `&#10003;` كنص له. مباشرة بعد ذلك العنصر `span` ولكن داخل عنصر القائمة، أضف النص `Ad-supported streaming`.

داخل عنصر القائمة الثاني، أنشئ عنصر `span` وأضف الكيان `&#10003;` كنص له. مباشرة بعد ذلك العنصر `span`، ولكن داخل عنصر القائمة، أضف النص `Curated playlists`.

# --hints--

يجب أن يكون لديك عنصران من عناصر القائمة داخل عنصر `ul` الخاص بك.

```js
assert.lengthOf(document.querySelectorAll("ul > li"), 2)
```

يجب أن يحتوي أول عنصر `li` على عنصر `span` يحتوي على كيان حرف HTML `&#10003;`.

```js
const listItems = document.querySelectorAll("ul > li")
const firstSpan = listItems[0].querySelector("span")

assert.equal(firstSpan?.textContent, "✓")
```

يجب أن يحتوي أول عنصر `li` على النص `Ad-supported streaming` مباشرة بعد عنصر `span`.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[0].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Ad-supported streaming");
```

يجب أن يحتوي ثاني عنصر `li` على عنصر `span` يحتوي على كيان حرف HTML `&#10003;`.

```js
const listItems = document.querySelectorAll("ul > li")
const secondSpan = listItems[1].querySelector("span")

assert.equal(secondSpan?.textContent, "✓")
```

يجب أن يحتوي ثاني عنصر `li` على النص `Curated playlists` مباشرة بعد عنصر `span`.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[1].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Curated playlists");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>Start exploring millions of songs with basic features and ads.</p>
            <ul>
            --fcc-editable-region--
              
            --fcc-editable-region--
            </ul>
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
