---
id: 6140cd32d018ed0f600eefce
title: Pasos 4
challengeType: 0
dashedName: step-4
---

# --description--

Establece la propiedad `position` del selector `.wheel` en `absolute`. Establece tanto la propiedad `height` como `width` en `55vw`.

# --hints--

El selector `.wheel` debe tener una propiedad `position` establecida en `absolute`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.position === 'absolute');
```

El selector `.wheel` debe tener una propiedad `height` establecida en `55vw`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.height === '55vw');
```

El selector `.wheel` debe tener una propiedad `width` establecida en `55vw`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.width === '55vw');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
}
--fcc-editable-region--
```
