---
id: 6140cbeec34e970dfe75e710
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Crea un selector para el elemento `wheel`. Inicia definiendo el `border` con los valores `2px solid black`, el `border-radius` a `50%` y el `margin-left` a `50px`.

# --hints--

Debes tener un selector `.wheel`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel'));
```

Tu selector `.wheel` debe tener una propiedad `border` establecida en `2px solid black`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.border === '2px solid black');
```

Tu selector `.wheel` debe tener una propiedad `border-radius` establecida en `50%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.borderRadius === '50%');
```

Tu selector `.wheel` debe tener una propiedad `margin-left` establecida en `50px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.marginLeft === '50px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
