---
id: 6140c7e645d8e905819f1dd4
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Per iniziare, aggiungi un elemento `link` per il file `./styles.css`.

# --hints--

Il tuo codice dovrebbe avere un elemento `link`.

```js
assert.match(code, /<link/);
```

Dovresti avere un elemento `link`.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

Il tuo elemento `link` dovrebbe essere all'interno del tuo elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Il tuo elemento `link` dovrebbe avere un attributo `rel` con il valore `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Il tuo elemento `link` dovrebbe avere un attributo `href` con il valore `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>

  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css

```
