---
id: 6140d0069049f5139d78da40
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

La proprietà `transform-origin` serve a impostare il punto attorno al quale viene applicata una trasformazione CSS. Per esempio, quando applichi una trasformazione `rotate` (come farai più avanti in questo progetto), `transform-origin` determina attorno a quale punto l’elemento viene ruotato.

Dai al selettore `.line` una proprietà `transform-origin` con valore `0% 0%`. Questo sposterà il punto di origine di `0%` da sinistra e `0%` dall’alto, impostandolo nell’angolo in alto a sinistra dell’elemento.

# --hints--

Il tuo selettore `.line` dovrebbe avere una proprietà `transform-origin` impostata su `0% 0%`.

```js
const transformOrigin = new __helpers.CSSHelp(document).getStyle('.line')?.transformOrigin;
assert.oneOf(transformOrigin, ['0% 0%', '0% 0% 0px']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
}
--fcc-editable-region--
```
