---
id: 617ace7d831f9c16a569b38a
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

Dai al tuo selettore `.wheel` una proprietà `max-height` e una `max-width` entrambe impostate su `500px`.

# --hints--

Il tuo selettore `.wheel` dovrebbe avere una proprietà `max-height` impostata su `500px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.maxHeight, '500px');
```

Il tuo selettore `.wheel` dovrebbe avere una proprietà `max-width` impostata su `500px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.maxWidth, '500px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
}
--fcc-editable-region--
```
