---
id: 66a37f37ef5823a313de8c26
title: Crie uma órbita lunar
challengeType: 25
dashedName: build-a-moon-orbit
demoType: onClick
---

# --description--

Neste laboratório, você criará uma animação simples da órbita da Lua ao redor da Terra usando HTML e CSS. A Terra estará no centro do sistema, enquanto a Lua girará ao redor dela.

Atenda às histórias de usuário abaixo e faça todos os testes passarem para completar o laboratório.

**Histórias de usuário:**

1. Você deve alinhar todos os elementos ao centro da página definindo a altura como `100%` da viewport e configurando um layout do flexbox no elemento `body`.

2. Você deve ter uma `div` com a classe `space`.

3. Dentro da `div` `.space`, deve haver dois outros elementos `div` com as classes `earth` e `orbit`, respectivamente, em ordem.

4. Dentro da `div` `.orbit`, deve haver outra `div` com a classe `moon`.

5. The `div` element with a `class` of `space` should be centered on the page and have a width and height of `200px`.

6. The `div` element with a `class` of `space` should use `relative` positioning.

7. The `.earth` element should use `absolute` positioning. Position the center of the `.earth` element at the halfway point of its parent on both the vertical (top) and horizontal (left) axes. After that, shift the `.earth` element back by half its own width and height, to center it within its parent, the `.space` element.

8. The `.earth` element should have a width and height of `100px`.

9.  The `.orbit` element should have a width and height of `200px`.

10.  The `.orbit` element should be positioned using `absolute` positioning. O canto inferior direito deve estar no centro da `div` `space` usando uma propriedade `transform`, que a desloque em `-50%` nos eixos vertical e horizontal.

11. O caminho da órbita para a lua ao redor da Terra deve ser um círculo.

12. The `.moon` element should be positioned using `absolute` positioning and have a width and height of `30px`. The `.moon` element should position itself at the top of the `orbit` `div` and be centered horizontally.

13. You should further adjust the horizontal positioning of the `.moon` element by moving the element to the left by half of its width.

14. Your `.earth` and `.moon` elements should have a background color and a `border-radius` of `50%` to make them look like planets.

15. Você deve definir uma animação `@keyframes orbit` que rotacione o elemento `.orbit` 360 graus em torno de seu centro. Você deve aplicar esta animação ao elemento `.orbit` com uma duração de `5` segundos, uma função de tempo linear e iterações infinitas.

**Observação:** certifique-se de vincular sua folha de estilo ao HTML e de aplicar o CSS.

# --hints--

O conteúdo do elemento `body` deve estar centralizado na página definindo a altura como `100%` da altura da viewport, `display` como `flex`, `justify-content` e `align-items` como `center`.

```js

assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.display, 'flex');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.justifyContent, 'center');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.alignItems, 'center');
```

Você deve ter uma `div` principal com a classe `space`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'space');
```

A `div` `space` deve ter uma largura e altura de `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.height, '200px');
```

The `.space` element should have a `position` property of `relative`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.position, 'relative')
```

You should have a `div` with a class of `earth` inside the `div` with a `class` of `space`.

```js
assert.exists(document.querySelector('div.space div.earth'));
```

A `div` `earth` deve ser o primeiro filho da `div` `space`.

```js
assert.strictEqual(document.querySelector('.space > div')?.getAttribute('class'), 'earth');
```

A `div` `earth` deve ter uma largura e altura de `100px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.width, '100px');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.height, '100px');
```

A `div` `earth` deve usar o posicionamento `absolute`.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.position, 'absolute');
```

A `div` `earth` deve ter posição `top` e `left` de `50%`.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.top, '50%') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.left, '50%');
```

A `div` `earth` deve ser movida para trás e para cima em `50%`, definindo a propriedade `transform` como `translate(-50%, -50%)`.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.transform, 'translate(-50%, -50%)');
```

Você deve ter uma `div` com a classe `orbit` dentro da `div` `.space`.

```js
assert.exists(document.querySelector('div.space div.orbit'));
```

A `div` de classe `orbit` deve vir após o elemento `div` de classe `earth`.

```js
assert.isTrue(document.querySelector('.earth')?.nextElementSibling?.classList?.contains('orbit'));
```

A `div` `orbit` deve ter uma largura e altura de `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.height, '200px');
```

A `div` `orbit` deve usar posicionamento `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.position, 'absolute');
```

A `div` `orbit` deve estar posicionada no centro da `div` `space` usando uma propriedade `transform` que a desloque em `-50%` nos dois eixos, vertical e horizontal.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.transform, 'translate(-50%, -50%)');
```

Você deve ter uma `div` com a classe `moon` dentro da `div` `.orbit`.

```js
assert.exists(document.querySelector('div.space div.orbit div.moon'));
```

A `div` `moon` deve ter uma largura e altura de `30px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.width, '30px') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.height, '30px');
```

A `div` `moon` deve ser posicionada usando posicionamento `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.position, 'absolute');
```

A borda superior do elemento `.moon` deve ser alinhada com a borda superior de seu bloco de contenção.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.top, '0px');
```

Você deve mover a borda esquerda do elemento `.moon` para o centro do bloco de contenção (`50%` do lado esquerdo).

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.left, '50%');
```

A `div` `.moon` deve ser ajustada horizontalmente em `-50%` no eixo x usando a propriedade translate.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.transform, 'translateX(-50%)');
```

O elemento `div` `.earth` deve ter uma cor de fundo.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.earth')?.backgroundColor);
```

O elemento `.earth` deve ter `border-radius` de `50%` para parecer um círculo.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.borderRadius, '50%');
```

O elemento `div` `.moon` deve ter uma cor de fundo.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.moon')?.backgroundColor);
```

O elemento `.moon` deve ter `border-radius` de `50%` para parecer um círculo.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.borderRadius, '50%');
```

Você deve criar uma tag `@keyframes`.

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('keyframes'), 1);
```

A nova regra `@keyframes` deve ser nomeada `orbit`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name, 'orbit');
```

A regra `@keyframes orbit` deve ter um seletor `0%`.

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules; 
assert.strictEqual(rules?.[0]?.keyText, '0%');
```

A regra `@keyframes orbit` deve ter um seletor `100%`.

```js 
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules
assert.isTrue(rules?.[0]?.keyText === '100%' || rules?.[1]?.keyText === '100%');
```

O seletor `0%` deve ter uma propriedade `transform` definida como `rotate(0deg) translate(-50%, -50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[0]?.style?.transform, 'rotate(0deg) translate(-50%, -50%)');
```

O elemento `100%` deve vir depois do elemento `0%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.keyText, '100%')
```

O seletor `100%` deve ter uma propriedade `transform` definida como `rotate(360deg) translate(-50%, -50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.style?.transform, 'rotate(360deg) translate(-50%, -50%)')
```

O seletor `.orbit` deve ter uma propriedade `animation` definida como `orbit 5s linear infinite`.

```js

const orbitStyles = new __helpers.CSSHelp(document).getStyle('.orbit');
assert.strictEqual(orbitStyles?.animationName, "orbit");
assert.strictEqual(orbitStyles?.animationDuration, "5s");
assert.strictEqual(orbitStyles?.animationTimingFunction, "linear");
assert.strictEqual(orbitStyles?.animationIterationCount, "infinite");

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Moon Orbit</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Moon Orbiting Earth</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="space">
        <div class="earth"></div>
        <div class="orbit">
            <div class="moon"></div>
        </div>
    </div>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #000;
}

.space {
    position: relative;
    width: 200px;
    height: 200px;
}

.earth {
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    left: 50%;
    transform: translate(-50%, -50%);
    box-shadow: 0 0 20px rgba(0, 0, 255, 0.5);
}

.orbit {
    width: 200px;
    height: 200px;
    position: absolute;
    transform: translate(-50%, -50%);
    animation: orbit 5s linear infinite;
}

.moon {
    width: 30px;
    height: 30px;
    background-color: gray;
    border-radius: 50%;
    position: absolute;
    top: 0;
    left: 50%;
    transform: translateX(-50%);
    box-shadow: 0 0 10px rgb(255, 255, 255);
}


@keyframes orbit {
    0% {
        transform: rotate(0deg) translate(-50%, -50%);
    }

    100% {
        transform: rotate(360deg) translate(-50%, -50%);
    }
}
```

