---
id: 60f027099a15b00485563dd2
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

Below the `DOCTYPE`, add an `html` element with a `lang` attribute set to `en`, so that you have a place to start putting some code.

# --hints--

يجب أن يكون تعريف `DOCTYPE` في بداية HTML الخاص بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

يجب أن يحتوي عنصرك `html` على علامة افتتاحية (opening tag) مع خاصية `lang` من `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصر `html` الخاص بك على علامة الأغلاق.

```js
assert(code.match(/<\/html\s*>/));
```

يجب أن تكون العلامات الخاصة بعنصر `html` في الترتيب الصحيح.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

يجب أن يكون لديك عنصر `html` واحد فقط.

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
