---
id: 642db8c409d9991d0b3b2f0d
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Il tuo progetto inizia con un contenitore HTML di base e del CSS corrispondente. La tua prima attività sarà generare programmaticamente le celle per il tuo foglio di calcolo.

L'oggetto globale `window` rappresenta la finestra del browser (o la scheda). Ha una proprietà `onload` che ti permette di definire il comportamento quando la finestra ha caricato l'intera pagina, inclusi fogli di stile e script.

Inizia impostando la proprietà `onload` di `window` su una funzione freccia senza parametri. All'interno della funzione, dichiara una variabile `container` e assegnale il valore ottenuto consultando l'elemento tramite l'id `id` di `"container"`.

# --hints--

Dovresti consultare la proprietà `onload` dell'oggetto `window`.

```js
assert.match(code, /window\.onload/);
```

Dovresti impostare la proprietà `onload` su una funzione.

```js
assert.isFunction(window.onload);
```

Dovresti usare la sintassi della funzione freccia.

```js
assert.match(code, /window\.onload\s*=\s*\(/);
```

La tua funzione `onload` non dovrebbe prendere parametri.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)/);
```

Dovresti dichiarare una variabile `container` nella tua funzione `onload`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*(?:let|var|const)\s+container/);
```

La tua variabile `container` dovrebbe essere dichiarata con `const`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container/);
```

Dovresti usare `document.getElementById()`

```js
assert.match(code, /document\.getElementById\(/);
```

Dovresti ottenere l'elemento con l'id `id` di `"container"`.

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

Dovresti assegnare l'elemento `#container` a `container`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
