---
id: 6434750c53db16218f41e6e1
title: Passaggio 15
challengeType: 0
dashedName: step-15
---

# --description--

Ricorda che `range()` restituisce un array, quindi puoi concatenare direttamente i metodi degli array alla chiamata della funzione.

Chiama `range()` con `1` e `99` come argomenti, e concatena il metodo `.forEach()`. Passa al metodo `.forEach()` una callback vuota che prende `number` come parametro.

# --hints--

Dovresti chiamare la tua funzione `range()`.

```js
assert.lengthOf(code.match(/range\(/g), 2);
```

Dovresti passare `1` come primo argomento alla tua chiamata `range()`.

```js
assert.match(code, /range\(\s*1/);
```

Dovresti passare `99` come secondo argomento alla tua chiamata `range()`.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)/);
```

Dovresti concatenare il metodo `.forEach()` alla tua chiamata `range()`.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(/);
```

Dovresti passare una funzione di callback a `.forEach()` usando la sintassi arrow.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

La tua funzione di callback dovrebbe avere `number` come unico parametro.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\(\s*number\s*\)|number)\s*=>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
--fcc-editable-region--

--fcc-editable-region--
}
```
