---
id: 61695ab9f6ffe951c16d03dd
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

قم بتسكين عنصر `head` داخل عنصر `html`. بعد عنصر `head` مباشرة، أضف عنصر `body`.

# --hints--

يجب أن يكون لديك وسم بداية `head`.

```js
assert(code.match(/<head\s*>/i));
```

يجب أن يكون لديك وسم الإغلاق `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

يجب أن يكون لديك وسم بداية `body`.

```js
assert(code.match(/<body\s*>/i));
```

يجب أن يكون لديك وسم الإغلاق `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

يجب أن يكون هناك وسم بداية `head` واحد فقط.

```js
assert(code.match(/<head\s*>/ig).length === 1);
```

يجب أن يكون هناك وسم الإغلاق `head` واحد فقط.

```js
assert(code.match(/<\/head\s*>/ig).length === 1);
```

يجب أن يكون هناك وسم بداية `body` واحد فقط.

```js
assert(code.match(/<body\s*>/ig).length === 1);
```

يجب أن يكون هناك وسم الإغلاق `body` واحد فقط.

```js
assert(code.match(/<\/body\s*>/ig).length === 1);
```

يجب أن يكون عنصر `head` فارغًا.

```js
assert(code.match(/<head\s*>\s*<\/head\s*>/i));
```

يجب أن يكون عنصر `body` فارغًا.

```js
assert(code.match(/<body\s*>\s*<\/body\s*>/i));
```

يجب أن يكون عنصر `body` موضوعًا بعد عنصر `head`.

```js
assert(code.match(/<\/head\s*>\s*<body\s*>/i));
```

يجب أن يكون عنصر `head` داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

يجب أن يكون عنصر `body` داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
