---
id: 61695197ac34f0407e339882
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

كما رأيت في المشاريع السابقة، يجب أن تبدأ صفحات الويب بـ `DOCTYPE html` إعلان، يليه عنصر `html`.

أضف إعلان `DOCTYPE html` في أعلى الوثيقة، وعنصر `html` بعد ذلك. أعطِ عنصر `html` خاصية `lang` بقيمة `en`.

# --hints--

يجب أن تحتوي شفرتك على مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

يجب أن تضع مسافة بعد مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

يجب أن تحدد نوع الوثيقة ليكون `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

يجب أن تغلق إعلان `DOCTYPE` بعلامة `>` بعد النوع.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

يجب أن يحتوي عنصر `html` على وسم بداية مع خاصية `lang` بقيمة `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصر `html` على وسم الإغلاق.

```js
assert(code.match(/<\/html\s*>/gi));
```

يجب أن يكون إعلان `DOCTYPE` في بداية شفرة HTML الخاصة بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
