---
id: 61695d1fbc003856628bf561
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

تحدد الخاصية `charset` ترميز الأحرف المستخدم في الوثيقة. `utf-8` (تنسيق تحويل يونيكود – 8-بت) هو معيار ترميز أحرف يُستخدم للاتصالات الإلكترونية.

داخل عنصر `head`، قم بتسكين عنصر `meta` مع الخاصية `charset` مضبوطة على القيمة `"utf-8"`.

# --hints--

يجب أن يكون لديك عنصر `meta` واحد.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 1);
```

يجب أن يكون عنصر `meta` عنصرًا فارغًا، ولا يجب أن يحتوي على وسم إغلاق `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

يجب أن يحتوي عنصر `meta` على خاصية `charset` مضبوطة على القيمة `"utf-8"`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

يجب أن يكون عنصر `meta` داخل عنصر `head`.

```js
const metaElementRegex = /<head\s*>(?:.|\r|\n)*?<meta(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, metaElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
