---
id: 616968c2c94c8071b349c146
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

الآن أنت جاهز لبدء إضافة المحتوى إلى الصفحة.

داخل عنصر `body`، قم بتسكين عنصر `h1` مع النص `CSS Color Markers`.

# --hints--

يجب أن تحتوي شفرتك على عنصر `h1`.

```js
const title = document.querySelector('h1');
assert.exists(title);
```

يجب أن يكون لديك وسم افتتاحي `h1`.

```js
assert(code.match(/<h1\s*>/i));
```

يجب أن يكون لديك وسم إغلاق `h1`.

```js
assert(code.match(/<\/h1\s*>/i));
```

يجب أن يكون عنصر `h1` داخل عنصر `body`.

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

يجب أن يحتوي عنصر `h1` على النص `CSS Color Markers`.

```js
const h1 = document.querySelector('h1');
assert.equal(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

تذكر، أن حالة الأحرف والتهجئة مهمة لنص `h1`.

```js
const h1 = document.querySelector('h1');
assert.equal(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
