---
id: 61695c4aad56f95497c19583
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

تذكر أن عنصر `title` يُعطي محركات البحث معلومات إضافية عن الصفحة. كما يعرض محتوى ذلك العنصر `title` بطريقتين إضافيتين:

* في شريط العنوان عندما تكون الصفحة مفتوحة  
* في تبويب المتصفح للصفحة عند تحريك المؤشر فوقه. حتى لو لم يكن ذلك التبويب نشطًا، بمجرد تحريك المؤشر فوقه، يُعرض نص `title`.

داخل عنصر `head`، قم بتسكين عنصر `title` يحتوي على النص `Colored Markers`.

# --hints--

يجب أن يكون لديك وسم `title` افتتاحي.

```js
assert(code.match(/<title\s*>/i));
```

يجب أن يكون لديك وسم `title` إغلاق.

```js
assert(code.match(/<\/title\s*>/i));
```

يجب أن يكون لمشروعك العنوان `Colored Markers`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

تذكر، أن حالة الأحرف والإملاء مهمان لنص العنوان.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
