---
id: 616965351e74d4689eb6de30
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

يمكنك إضافة عدة عناصر `meta` في صفحة ويب. كل عنصر `meta` يضيف معلومات عن الصفحة لا يمكن التعبير عنها بعناصر HTML الأخرى.

أضف عنصر `meta` آخر داخل `head`. عرّف له خاصية `name` بقيمة `"viewport"` وخاصية `content` بقيمة `"width=device-width, initial-scale=1.0"` لكي تبدو صفحتك نفسها على جميع الأجهزة.

# --hints--

يجب أن يكون لديك عنصران `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 2);
```

يجب أن يكون عنصر `meta` الجديد عنصرًا فارغًا، فلا يجب أن يحتوي على وسم الإغلاق `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```


يجب أن يحتوي عنصر `meta` الجديد على خاصية `name` بقيمة `"viewport"`، وخاصية `content` بقيمة `"width=device-width, initial-scale=1.0"`.

```js
const meta = [...document.querySelectorAll('meta')];
const contentAttrRegex = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/i;
const target = meta?.find(m => m?.getAttribute('name')?.toLowerCase()?.replace(/\s*/g, '') === 'viewport' && contentAttrRegex.test(m?.getAttribute('content')));
assert.exists(target);
```

يجب أن يكون العنصران `meta` داخل عنصر `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('meta').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
