---
id: 616971b3cd990186b66c99a1
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

الآن بعد أن أعددت ملف CSS الخارجي، يمكنك البدء في تنسيق الصفحة.

كتذكير، إليك كيفية استهداف عنصر الفقرة ومحاذاته إلى اليمين:

```css
p {
  text-align: right;
}
```

أنشئ قاعدة CSS جديدة تستهدف عنصر `h1`، واضبط خاصية `text-align` على `center`.

# --hints--

يجب أن تستخدم محدد `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

يجب أن تحتوي قاعدة CSS الخاصة بـ `h1` على خاصية `text-align` مضبوطة على `center`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
