---
id: 616d4a84b756d9c4b8255093
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

حان الوقت لإضافة بعض اللون إلى العلامة. تذكر أن إحدى الطرق لإضافة لون إلى عنصر هي استخدام <dfn>كلمة لون مفتاحية</dfn> مثل `black`، `cyan`، أو `yellow`.

كتذكير، إليك كيفية استهداف الصنف `freecodecamp`:

```css
.freecodecamp {
  
}
```

أنشئ قاعدة CSS جديدة تستهدف الصنف `marker`، واضبط خاصية `background-color` على `red`.

**ملاحظة:** لن ترى أي تغييرات بعد إضافة CSS.

# --hints--

يجب أن تنشئ محدِّد صنف لاستهداف الصنف `marker`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker'));
```

يجب أن تحتوي قاعدة CSS الخاصة بـ `.marker` على خاصية `background-color` مضبوطة على `red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor === 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
