---
id: 616d55bd160a95e22453a081
title: الخطوة 14
challengeType: 0
dashedName: step-14
---

# --description--

الآن بعد أن أصبح لديك علامة واحدة في الوسط مع لون، حان الوقت لإضافة العلامات الأخرى.

في `div` الحاوية `container`، أضف عنصرين `div` آخرين وأعط كلًا منهما صنفًا `marker`.

# --hints--

يجب أن يحتوي أول عنصر `div` جديد على وسم بداية.

```js
assert([...code.matchAll(/<div.*?>/gi)][2]);
```

يجب أن يحتوي أول عنصر `div` جديد على وسم الإغلاق.

```js
assert([...code.matchAll(/<\/div\s*>/gi)][2]);
```

يجب أن يحتوي ثاني عنصر `div` جديد على وسم بداية.

```js
assert([...code.matchAll(/<div.*?>/gi)][3]);
```

يجب أن يحتوي ثاني عنصر `div` جديد على وسم الإغلاق.

```js
assert([...code.matchAll(/<\/div\s*>/gi)][3]);
```

يجب أن تكون عناصر `div` الجديدة داخل `div` الذي يحمل الصنف `container`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.localName === 'div') && containerChildren.length === 3);
```

يجب أن يكون لكل من عناصر `div` الجديدة صنف `marker`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    <div class="container">
      <div class="marker">
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
