---
id: 616d525007b8c5d8b3308b1c
title: الخطوة 13
challengeType: 0
dashedName: step-13
---

# --description--

سيبدو مؤشرُك أفضل إذا كان في وسط الصفحة. طريقة سهلة لتحقيق ذلك هي باستخدام الخاصية المختصرة `margin` <dfn>خاصية مختصرة</dfn>.

في المشروع السابق، ضبطت مساحة الهامش للعناصر بشكل منفصل باستخدام خصائص مثل `margin-top` و`margin-left`. الخاصية المختصرة `margin` تُسهّل ضبط عدة مساحات هامش في نفس الوقت.

لتمركز المؤشر على الصفحة، اضبط خاصية `margin` له على `auto`. هذا يضبط `margin-top` و`margin-right` و`margin-bottom` و`margin-left` جميعها على `auto`.

# --hints--

يجب أن تحتوي قاعدة CSS الخاصة بـ `.marker` على خاصية `margin` مضبوطة على `auto`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
}
--fcc-editable-region--

```
