---
id: 6176519636a76b810ab1219a
title: الخطوة 20
challengeType: 0
dashedName: step-20
---

# --description--

أنشئ قاعدة CSS تستهدف الصنف `two` واضبط خاصية `background-color` على `green`.

أنشئ أيضًا قاعدة CSS منفصلة تستهدف الصنف `three` واضبط خاصية `background-color` على `blue`.

# --hints--

يجب أن تستخدم محدد صنف لاستهداف الصنف `two`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two'));
```

يجب أن تحتوي قاعدة CSS الخاصة بـ `.two` على خاصية `background-color` مضبوطة على `green`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'green');
```

يجب أن تستخدم محدد صنف لاستهداف الصنف `three`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three'));
```

يجب أن تحتوي قاعدة CSS الخاصة بـ `.three` على خاصية `background-color` مضبوطة على `blue`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

--fcc-editable-region--

--fcc-editable-region--

```
