---
id: 61764c602bee6974e7790f35
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

لإعطاء العلامات ألوانًا مختلفة، تحتاج إلى إضافة صنف فريد لكل منها. يمكن إضافة عدة أصناف إلى عنصر عن طريق سردها في الخاصية `class` وفصلها بمسافة. على سبيل المثال، يضيف التالي صنفي `animal` و `dog` إلى عنصر `div`.

```html
<div class="animal dog">
```

إذا أضفت عدة أصناف إلى عنصر HTML، قد تُلغى أنماط الأصناف الأولى التي تذكرها بواسطة الأصناف اللاحقة.

لبداية، أضف الصنف `one` إلى أول عنصر `div` للعلامة.


# --hints--

يجب أن تضيف الصنف `one` إلى أول عنصر `div` للعلامة.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

يجب أن يحتوي أول عنصر `div` للعلامة على الأصناف `marker` و `one`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
