---
id: 617b65579ce424bf5f02ca73
title: الخطوة 21
challengeType: 0
dashedName: step-21
---

# --description--

هناك نموذجان رئيسيان للألوان: نموذج RGB (أحمر، أخضر، أزرق) <dfn>الإضافي</dfn> المستخدم في الأجهزة الإلكترونية، ونموذج CMYK (سماوي، أرجواني، أصفر، أسود) <dfn>الطرحي</dfn> المستخدم في الطباعة.

في هذا المشروع، ستعمل مع نموذج RGB. هذا يعني أن الألوان تبدأ باللون الأسود، وتتغير مع إدخال مستويات مختلفة من الأحمر والأخضر والأزرق. طريقة سهلة لرؤية ذلك هي باستخدام دالة `rgb` في CSS.

أنشئ قاعدة CSS جديدة تستهدف الصنف `container` واضبط خاصية `background-color` إلى الأسود باستخدام `rgb(0, 0, 0)`.

# --hints--

يجب أن تستخدم محدد صنف لاستهداف الصنف `container`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

يجب أن تحتوي قاعدة CSS الخاصة بـ `.container` على خاصية `background-color` مضبوطة على `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
