---
id: 617bc3386dc7d07d6469bf20
title: الخطوة 32
challengeType: 0
dashedName: step-32
---

# --description--

لاحظ أنه لإنشاء اللون البرتقالي، كان عليك زيادة شدة الأحمر وتقليل شدة قيم `rgb` الخاصة بالأخضر. ذلك لأن البرتقالي هو مزيج من الأحمر والأصفر.

لإنشاء اللون الثالثي الأخضر الربيعي، اجمع بين السماوي والأخضر. حدّث دالة `rgb` في قاعدة CSS الخاصة بـ `.two` بحيث يكون الأخضر عند القيمة القصوى، واضبط الأزرق على `127`.

# --hints--

يجب أن تحتوي قاعدة CSS الخاصة بـ `.two` على خاصية `background-color` مضبوطة على `rgb(0, 255, 127)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 255, 255);
}
--fcc-editable-region--

.three {
  background-color: rgb(255, 0, 255);
}

```
