---
id: 618a00ed1ca871a2b3aca0cb
title: الخطوة 35
challengeType: 0
dashedName: step-35
---

# --description--

الآن بعد أن استعرضت جميع الألوان الأولية والثانوية والثالثية على عجلة الألوان، سيكون من الأسهل فهم مفاهيم نظرية الألوان الأخرى وكيف تؤثر على التصميم.

أولًا، في قواعد CSS `.one` و`.two` و`.three`، عدل القيم في دالة `rgb` بحيث تُضبط خاصية `background-color` لكل عنصر على اللون الأسود النقي. تذكر أن دالة `rgb` تستخدم نموذج الألوان الإضافي، حيث تبدأ الألوان بالأسود وتتغير مع زيادة قيم الأحمر والأخضر والأزرق.

# --hints--

يجب أن تحتوي قاعدة CSS `.one` على خاصية `background-color` مضبوطة على `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(0, 0, 0)');
```

يجب أن تحتوي قاعدة CSS `.two` على خاصية `background-color` مضبوطة على `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

يجب أن تحتوي قاعدة CSS `.three` على خاصية `background-color` مضبوطة على `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(127, 255, 0);
}

.two {
  background-color: rgb(0, 127, 255);
}

.three {
  background-color: rgb(255, 0, 127);
}
--fcc-editable-region--

```
