---
id: 619b7fd56aa2253778bcf5f7
title: الخطوة 48
challengeType: 0
dashedName: step-48
---

# --description--

نموذج الألوان <dfn>HSL</dfn>، أو اللون، التشبع، وفاتح اللون، هو طريقة أخرى لتمثيل الألوان.

تقبل دالة CSS `hsl` ثلاث قيم: رقم من 0 إلى 360 للون، ونسبة مئوية من 0 إلى 100 للتشبع، ونسبة مئوية من 0 إلى 100 لفاتح اللون.

إذا تخيلت عجلة الألوان، فاللون الأحمر عند 0 درجة، والأخضر عند 120 درجة، والأزرق عند 240 درجة.

التشبع هو شدة اللون من 0%، أي الرمادي، إلى 100% للون النقي. يجب إضافة علامة النسبة المئوية `%` لقيم التشبع وفاتح اللون.

فاتح اللون هو مدى سطوع اللون، من 0%، أي الأسود الكامل، إلى 100%، الأبيض الكامل، مع كون 50% محايدًا.

في قاعدة CSS `.blue`، استخدم دالة `hsl` لتغيير خاصية `background-color` إلى اللون الأزرق النقي. اضبط اللون على `240`، والتشبع على `100%`، وفاتح اللون على `50%`.

# --hints--

يجب أن تحتوي قاعدة CSS `.blue` على خاصية `background-color` مضبوطة على `hsl(240, 100%, 50%)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: #007F00;
}

--fcc-editable-region--
.blue {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
