---
id: 61a493ead935c148d2b76312
title: الخطوة 50
challengeType: 0
dashedName: step-50
---

# --description--

دالة `linear-gradient` مرنة جدًا — هذه هي الصيغة النحوية الأساسية التي ستستخدمها في هذا الدرس:

```css
linear-gradient(gradientDirection, color1, color2, ...);
```

`gradientDirection` هو اتجاه الخط المستخدم للانتقال. `color1` و`color2` هما معلمتا لون، وهما الألوان التي ستُستخدم في الانتقال نفسه. يمكن أن تكون هذه أي نوع من الألوان، بما في ذلك كلمات لون مفتاحية، أو ألوان سداسية، أو `rgb`، أو `hsl`.

الآن ستطبق تدرجًا لونيًا من الأحمر إلى الأخضر على طول خط بزاوية 90 درجة على العلامة الأولى.

أولًا، في قاعدة CSS الخاصة بـ `.red`، عرّف خاصية `background` لتكون `linear-gradient()`، ومرر لها القيمة `90deg` كـ `gradientDirection`.

# --hints--

يجب أن تحتوي قاعدة CSS الخاصة بـ `.red` على خاصية `background` بقيمة `linear-gradient(90deg)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
