---
id: 6995d0325fd7da40854fac40
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In einer vorherigen Lektion haben Sie die binäre Suche kennengelernt und erfahren, wie sie sich von der linearen Suche unterscheidet.

In diesem Workshop setzen Sie das in die Praxis um, indem Sie einen Algorithmus für die binäre Suche erstellen.

Beginnen Sie damit, eine `binarySearch`-Funktion mit `searchList`- und `value`-Parametern zu erstellen.

`searchList` ist das Array von Elementen, in dem die Funktion suchen wird, und `value` ist das Ziel, das die Funktion im `searchList` finden soll.

# --hints--

Sie sollten eine `binarySearch`-Funktion erstellen.

```js
assert.isFunction(binarySearch);
```

Ihre `binarySearch`-Funktion sollte `searchList`- und `value`-Parameter haben.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
assert.lengthOf(binarySearch?.parameters, 2);
assert.equal(binarySearch?.parameters[0].toString(), "searchList");
assert.equal(binarySearch?.parameters[1].toString(), "value");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
