---
id: 6996c95e5bede03bdf2937a2
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Als Nächstes erstellen Sie eine Bedingung, die überprüft, ob das Ziel `value` in der Mitte ist.

Erstellen Sie innerhalb der `while`-Schleife eine `if`-Anweisung, die prüft, ob das Ziel `value` gleich `valueAtMiddle` ist. Wenn ja, geben Sie die Variable `pathToTarget` zurück.

# --hints--

Sie sollten eine `if`-Anweisung mit der Bedingung `value === valueAtMiddle` erstellen.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)/
);
```

Sie sollten die Variable `pathToTarget` innerhalb der `if`-Anweisung zurückgeben.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?return\s+pathToTarget/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
