---
id: 6996ca21a74104721f2c7e5d
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

Sie müssen Dinge testen, damit Sie den Ablauf des Algorithmus in diesem Anfangsstadium verstehen.

Dafür müssen Sie zuerst aus der Schleife ausbrechen. Das liegt daran, dass die aktuelle Implementierung nur eine Iteration zulässt, sodass es bei Nichterfüllung der Bedingung zu einer Endlosschleife kommt.

Direkt nach der `if`-Anweisung verwenden Sie das Schlüsselwort `break`, um aus der `while`-Schleife auszubrechen. Danach geben Sie nach der `while`-Schleife ein leeres Array zurück, um anzuzeigen, dass der Wert nicht gefunden wurde.

# --hints--

Sie sollten das Schlüsselwort `break` verwenden, um nach der `while`-Anweisung aus der `if`-Schleife auszubrechen.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?\}\s*break\s*;?/
);
```

Sie sollten nach Ihrer `while`-Schleife ein leeres Array zurückgeben.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /return\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--
    
  }
  
--fcc-editable-region--
}
```
