---
id: 6996cfb6c01e9a802add38be
title: Schritt 13
challengeType: 1
dashedName: step-13
---

# --description--

Schließlich, wenn der `value` weder in der Mitte gefunden wird noch größer als `valueAtMiddle` ist, dann muss `value` kleiner als `valueAtMiddle` sein und sich links befinden. Der abschließende `else` Block wird dies behandeln.

Beenden Sie also die Schleife, indem Sie das `break` Schlüsselwort entfernen und einen `else` Block hinzufügen. Aktualisieren Sie darin die `high` Variable, indem Sie `1` von `mid` subtrahieren.

# --hints--

Sie sollten das `break` Schlüsselwort durch einen `else` Block ersetzen, der die `high` Variable auf `mid - 1` aktualisiert.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s*\{\s*high\s*=\s*mid\s*-\s*1\s*;?\s*\}/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
--fcc-editable-region--      
    } 
    break;
--fcc-editable-region--
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
