---
id: 6996c8f94fe6b250f11a2a1e
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Fügen Sie nun `valueAtMiddle` in das `pathToTarget`-Array ein, um die während der Suche durchgeführten Schritte zu verfolgen, unabhängig davon, ob es sich um das Ziel handelt. Das können Sie mit der `push()`-Methode erledigen.

# --hints--

Sie sollten die `push()`-Methode verwenden, um `valueAtMiddle` zum `pathToTarget`-Array hinzuzufügen.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /pathToTarget\s*\.\s*push\(\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
