---
id: 6996c796385b6ea4e222d5e8
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Um mit der Schleife zu beginnen, müssen Sie den mittleren Index des aktuellen Suchbereichs finden.

Sie können den mittleren Index berechnen, indem Sie den Durchschnitt der Variablen `low` und `high` ermitteln. Sie müssen eine Ganzzahldivision verwenden, um den Durchschnitt nach dem Addieren der beiden Werte zu erhalten, sodass das Ergebnis immer eine nach unten gerundete ganze Zahl ist, da die Indizes ganze Zahlen sein müssen.

Erstellen Sie innerhalb Ihrer `while`-Schleife eine `mid`-Variable, die auf den Durchschnitt von `low` und `high` gesetzt ist.

# --hints--

Innerhalb der while-Schleife sollten Sie eine `mid`-Variable erstellen, die auf `Math.floor((low + high) / 2)` gesetzt ist.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+mid\s*=\s*Math\.floor\(\s*\(\s*low\s*\+\s*high\s*\)\s*\/\s*2\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
--fcc-editable-region--
      
--fcc-editable-region--
  }
}
```
