---
id: 6996d1032feed88e2957dc8e
title: Schritt 15
challengeType: 1
dashedName: step-15
---

# --description--

Sie sehen, dass der letzte Aufruf ein leeres Array ausgibt. Um genauer anzugeben, dass der Wert nicht im gesuchten Array enthalten ist, aktualisieren Sie die `return`-Anweisung so, dass sie ein Array mit zwei Elementen zurückgibt: ein leeres Array und den String `Value not found`.

# --hints--

Sie sollten die `return`-Anweisung aktualisieren, sodass die Funktion ein Array zurückgibt, bei dem `[]` das erste Element und `"Value not found"` das zweite Element ist.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const validExpressions = [
  "[[], 'Value not found']",
  "[[], \"Value not found\"]",
  "[[], `Value not found`]"
];
assert.isTrue(validExpressions.some(expr => binarySearch?.hasReturn(expr) || binarySearch?.hasReturn("(" + expr + ")")));
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
--fcc-editable-region--
  return [];
--fcc-editable-region--
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```
