---
id: 6996d0a2c2ac55041145a13f
title: Schritt 14
challengeType: 1
dashedName: step-14
---

# --description--

Ist Ihnen aufgefallen, dass der zweite Funktionsaufruf jetzt die Werte `3, 5, 4` in der Liste enthält? Das zeigt, dass die binäre Suche wie vorgesehen funktioniert.

So ist es passiert: Der Algorithmus hat zuerst `3` als Mitte der ursprünglichen Liste überprüft. Da `4` größer ist als `3`, verlagerte sich die Suche auf die rechte Hälfte. Dann wurde `5` als neue Mitte untersucht. Weil `4` kleiner ist als `5`, ging die Suche nach links weiter und identifizierte schließlich `4` als Mitte des letzten Bereichs.

Um die Funktion erneut zu testen, rufen Sie sie mit `[1, 3, 5, 9, 14, 22], 10` auf und geben Sie das Ergebnis in der Konsole aus. Dies ist eine Situation, in der der Wert nicht gefunden wird.

# --hints--

Sie sollten die Funktion mit `binarySearch([1, 3, 5, 9, 14, 22], 10)` aufrufen und das Ergebnis in der Konsole ausgeben.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*3\s*,\s*5\s*,\s*9\s*,\s*14\s*,\s*22\s*\]\s*,\s*10\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
--fcc-editable-region--

--fcc-editable-region--
```
