---
id: 6996c8375f4301004459ee85
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Nachdem Sie den `mid`-Index berechnet haben, müssen Sie den Wert aus dem `searchList` an diesem Index tatsächlich abrufen. Dieser Wert ist derjenige, den Sie mit Ihrem `value`-Parameter vergleichen, der das Ziel Ihrer Suche ist.

Erstellen Sie eine `valueAtMiddle`-Variable und weisen Sie ihr das Element aus `searchList` zu, das sich am `mid`-Index befindet.

# --hints--

Sie sollten eine `valueAtMiddle`-Variable erstellen, die auf `searchList[mid]` gesetzt ist.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+valueAtMiddle\s*=\s*searchList\s*\[\s*mid\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
