---
id: 6996c7363a4a437839f1f15d
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Nachdem Sie die Grenzen Ihrer Suche definiert haben, ist es Zeit, die Schleife zu erstellen, die den binären Suchalgorithmus ausführt.

Die binäre Suche funktioniert, indem sie den Suchbereich wiederholt eingrenzt. Dieser Vorgang läuft so lange weiter, wie ein gültiger Bereich von Elementen zum Überprüfen vorhanden ist.

Sie können dies mit einer `while`-Schleife ausdrücken, die so lange fortgesetzt wird, wie Ihr `low`-Zeiger kleiner oder gleich Ihrem `high`-Zeiger ist.

Wenn `low` jemals größer als `high` wird, bedeutet das, dass der Suchbereich leer geworden ist und das `value` nicht im Array enthalten ist.

Erstellen Sie innerhalb der `binarySearch`-Funktion eine `while`-Schleife mit einer Bedingung, die überprüft, ob `low` kleiner oder gleich high ist.

# --hints--

Sie sollten eine `while`-Schleife erstellen, die läuft, solange `low` kleiner oder gleich `high` ist.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /while\s*\(\s*low\s*<=\s*high\s*\)/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
