---
id: 6996c60281feca6718cc15c1
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Erstellen Sie innerhalb Ihrer `binarySearch`-Funktion eine `pathToTarget`-Variable und initialisieren Sie sie zunächst mit einem leeren Array.

# --hints--

Sie sollten innerhalb Ihrer `binarySearch`-Funktion eine Variable `pathToTarget` deklarieren und mit einem leeren Array initialisieren.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { pathToTarget } = binarySearch?.variables ?? {};
const value = eval(pathToTarget?.value.toString())
assert.isArray(value);
assert.isEmpty(value);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
