---
id: 6996c60281feca6718cc15c1
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Erstellen Sie innerhalb Ihrer `binarySearch`-Funktion eine `pathToTarget`-Variable und initialisieren Sie sie zunächst mit einem leeren Array.

# --hints--

Sie sollten eine `pathToTarget`-Variable haben, die innerhalb Ihrer `[]`-Funktion auf ein leeres Array `binarySearch` gesetzt ist.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+pathToTarget\s*=\s*\[\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
