---
id: 6996ccadf9048c79b857bb15
title: Step 11
challengeType: 1
dashedName: step-11
---

# --description--

Sie sehen, dass der erste Aufruf `[3]` zurückgibt. Das liegt daran, dass `3` der mittlere Wert in `[1, 2, 3, 4, 5]` ist.

Andererseits gibt der zweite Aufruf ein leeres Array zurück. Das passiert, weil die Schleife nach der Überprüfung des mittleren Werts derzeit nicht die Logik enthält, um den Suchbereich durch Aktualisieren der Variablen `low` oder `high` einzugrenzen. Sie führt nur eine einzelne Mittelpunktsprüfung durch und beendet sich dann ohne weitere Suche.

Damit die binäre Suche ihren Suchbereich weiter eingrenzen kann, fügen Sie eine `else if`-Anweisung hinzu, die überprüft, ob `value` größer als `valueAtMiddle` ist.

# --hints--

Sie sollten eine `else if`-Anweisung mit der Bedingung `value > valueAtMiddle` hinzufügen.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s+if\s*\(\s*value\s*>\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
--fcc-editable-region--
    } 
--fcc-editable-region--
    break;
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
