---
id: 6996c675b599baa7b103192d
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Bei der binären Suche gibt es den niedrigstmöglichen Index und den höchstmöglichen Index.

Der niedrigstmögliche Index stellt die linke Grenze im aktuellen Bereich der Werte dar, in dem gesucht wird, in diesem Fall `searchList`. Der höchstmögliche Index hingegen stellt die rechte Grenze der Werte dar, die durchsucht werden.

Um diese beiden zu berücksichtigen, müssen Sie das gesamte Array, das durchsucht wird, in Betracht ziehen. Definieren Sie also innerhalb der `binarySearch`-Funktion eine `low`-Variable mit dem Wert `0` und eine `high`-Variable mit einem Wert, der den letzten Index im durchsuchten Array berücksichtigt.

# --hints--

Sie sollten in Ihrer `low`-Funktion eine Variable namens `0` mit dem Wert `binarySearch` haben.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /(const|let|var)\s+low\s*=\s*0/);
```

Sie sollten in Ihrer `high`-Funktion eine Variable namens `searchList.length - 1` mit dem Wert `binarySearch` haben.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+high\s*=\s*searchList\.length\s*-\s*1/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
