---
id: 6996c675b599baa7b103192d
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Bei der binären Suche gibt es den niedrigstmöglichen Index und den höchstmöglichen Index.

Der niedrigstmögliche Index stellt die linke Grenze im aktuellen Bereich der Werte dar, in dem gesucht wird, in diesem Fall `searchList`. Der höchstmögliche Index hingegen stellt die rechte Grenze der Werte dar, die durchsucht werden.

Um diese beiden zu berücksichtigen, müssen Sie das gesamte Array betrachten, das durchsucht wird. Deklarieren Sie also innerhalb der `binarySearch`-Funktion eine Variable `low` mit dem Wert `0` und eine Variable `high` mit einem Wert, der den letzten Index im durchsuchen Array berücksichtigt.

# --hints--

Sie sollten in Ihrer `binarySearch`-Funktion eine Variable namens `low` mit dem Wert `0` deklarieren.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { low } = binarySearch?.variables ?? {};
assert.isTrue(low?.value.matches("0"));
```

Sie sollten in Ihrer `binarySearch`-Funktion eine Variable namens `high` mit dem Wert `searchList.length - 1` deklarieren.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { high } = binarySearch?.variables ?? {};
assert.isTrue(high?.value.matches("searchList.length - 1"));
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
