---
id: 6996d1032feed88e2957dc8e
title: Step 15
challengeType: 1
dashedName: step-15
---

# --description--

Sie sehen, dass der letzte Aufruf ein leeres Array ausgibt. Um weiter zu spezifizieren, dass der Wert nicht im gesuchten Array enthalten ist, aktualisieren Sie die `return`-Anweisung so, dass sie ein Array mit zwei Elementen zurückgibt: ein leeres Array und den String `Value not found`.

# --hints--

Sie sollten die `return`-Anweisung aktualisieren, um `[], 'Value not found'` zurückzugeben.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /return\s*\[\s*\[\s*\]\s*,\s*('|"|`)Value not found\1\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
--fcc-editable-region--
  return [];
--fcc-editable-region--
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```
