---
id: 6996cb5869acf58de751ae81
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Erinnern Sie sich, dass die Funktion ein Array namens `searchList` und das `value` entgegennimmt, nach dem Sie suchen.

Rufen Sie nun die Funktion mit `binarySearch([1, 2, 3, 4, 5], 3)` und `binarySearch([1, 2, 3, 4, 5, 9], 4)` auf und geben Sie beide Anweisungen in der Konsole aus.

# --hints--

Sie sollten die Funktion mit `binarySearch([1, 2, 3, 4, 5], 3)` aufrufen und das Ergebnis in der Konsole ausgeben.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(binarySearch([1,2,3,4,5],3))";
assert.include(strippedCode, expectedStatement);
```

Sie sollten die Funktion mit `binarySearch([1, 2, 3, 4, 5, 9], 4)` aufrufen und das Ergebnis in der Konsole ausgeben.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(binarySearch([1,2,3,4,5,9],4))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;

  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    }
    break;
  }
  return [];
}

--fcc-editable-region--

--fcc-editable-region--
```
