---
id: 6976db7ecfa770bf21307b20
title: プレイリストリミックスエンジンを作る
challengeType: 26
dashedName: build-a-playlist-remix-engine
---

# --description--

このラボでは、リスナーから提出された複数のプレイリストから1つのリミックスプレイリストを作成するプログラムを作ります。

各リスナーは聴きたい曲のリストを提供します。曲が複数回登場することもあり、同じアーティストが何度も現れることもあります。これらのプレイリストをステップごとに処理してください：1つのリストにまとめ、各曲にスコアを付け、重複する曲を削除し、同じアーティストが登場する回数を制限し、最終的な再生順を作成します。

**目的：** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. `flattenPlaylists` という名前の関数を作成してください。この関数は、各プレイリストが `trackId`、`artist`、`title`、`votes`、`bpm` のプロパティを持つオブジェクトの配列である配列を受け取ります。入力が配列でない場合、`flattenPlaylists` は空の配列を返す必要があります。例としてプレイリストが提供されていますので、関数のテストに使ってください。

2. `flattenPlaylists` は、元のトラックのすべてのプロパティに加え、トラックの出所を示すプレイリストのインデックスとトラックのインデックスを含む配列をセットした `source` プロパティを持つトラックオブジェクトのフラットな配列を返す必要があります。

3. `scoreTracks` という名前の関数を作成してください。この関数は `flattenPlaylists` が返すトラックオブジェクトの配列（各オブジェクトは `trackId`、`artist`、`title`、`votes`、`bpm`、`source` のプロパティを持つ）を受け取り、`score` プロパティを追加した新しいトラックオブジェクトの配列を返します。`score` は次の式で計算されます：`votes * 10 - Math.abs(bpm - 120)`。

4. `dedupeTracks` という名前の関数を作成してください。この関数は `scoreTracks` が返すトラックオブジェクトの配列を受け取り、重複する `trackId` のエントリーを削除し、最初の出現のみを保持した新しい配列を返します。

5. `enforceArtistQuota` という名前の関数を作成してください。この関数は `dedupeTracks` が返すトラックオブジェクトの配列と、アーティストごとの最大許容出現回数を表す数値を受け取ります。関数は、指定された回数を超えて同じアーティストが登場しないようにし、最も早い出現を保持した新しい配列を返します。

6. `buildSchedule` という名前の関数を作成してください。この関数は `enforceArtistQuota` が返すトラックオブジェクトの配列を受け取り、放送順での各トラックの位置を表す1始まりのインデックス `slot` を持つ `{ slot, trackId }` オブジェクトの新しい配列を返します。

7. `remixPlaylist` という名前の関数を作成してください。この関数はプレイリストの配列とアーティストごとの最大許容出現回数を受け取り、`flattenPlaylists`、`scoreTracks`、`dedupeTracks`、`enforceArtistQuota`、`buildSchedule` を順に呼び出して、最終的な放送スケジュールを `{ slot, trackId }` オブジェクトの配列として返します。

# --hints--

`flattenPlaylists` という名前の関数を作成してください。

```js
assert.isFunction(flattenPlaylists);
```

入力が配列でない場合、`flattenPlaylists` から空の配列を返す必要があります。

```js
assert.deepEqual(flattenPlaylists(null), []);
assert.deepEqual(flattenPlaylists(undefined), []);
assert.deepEqual(flattenPlaylists({}), []);
assert.deepEqual(flattenPlaylists("not an array"), []);
assert.deepEqual(flattenPlaylists(123), []);
```

`flattenPlaylists` が返す各トラックには、プレイリストのインデックスとトラックのインデックスを含む配列である `source` フィールドが含まれている必要があります。

```js
const playlists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 110 }
  ],
  [{ trackId: "t3", artist: "C", title: "Song 3", votes: 5, bpm: 130 }]
];

const flat = flattenPlaylists(playlists);

assert.isArray(flat);
assert.lengthOf(flat, 3);

flat.forEach((t) => {
  assert.property(t, "source");
  assert.isArray(t.source);
  assert.lengthOf(t.source, 2);
  assert.isNumber(t.source[0]);
  assert.isNumber(t.source[1]);
});

assert.deepEqual(flat[0].source, [0, 0]);
assert.deepEqual(flat[1].source, [0, 1]);
assert.deepEqual(flat[2].source, [1, 0]);
```

`scoreTracks` という名前の関数を作成してください。

```js
assert.isFunction(scoreTracks)
```

`scoreTracks` が返す各トラックには数値の `score` フィールドが含まれている必要があります。

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120, source: [0, 0] }
];

const scored = scoreTracks(tracks);

assert.isArray(scored);
assert.lengthOf(scored, 1);

assert.property(scored[0], "score");
assert.isNumber(scored[0].score);
```

目標BPMを `120` とし、次の式で `score` を計算してください：`votes * 10 - Math.abs(bpm - 120)`。

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 3, bpm: 120, source: [0, 0] }, // 30
  { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 100, source: [0, 1] },  // -10
  { trackId: "t3", artist: "C", title: "Song 3", votes: 2, bpm: 135, source: [1, 0] }  // 5
];

const formulaScored = scoreTracks(tracks);

const s1 = formulaScored.find(t => t.trackId === "t1").score;
const s2 = formulaScored.find(t => t.trackId === "t2").score;
const s3 = formulaScored.find(t => t.trackId === "t3").score;

assert.equal(s1, 30);
assert.equal(s2, -10);
assert.equal(s3, 5);
```

`dedupeTracks` という名前の関数を作成してください。

```js
assert.isFunction(dedupeTracks)
```

重複する `trackId` の値がある場合、`dedupeTracks` は最初の出現のみを保持する必要があります。

```js
const dupes = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "B", title: "Song 2", votes: 2, bpm: 119, source: [0, 1], score: 19 },
  { trackId: "t1", artist: "A", title: "Song 3", votes: 9, bpm: 140, source: [1, 0], score: 70 }
];

const deduped = dedupeTracks(dupes);

assert.isArray(deduped);
assert.lengthOf(deduped, 2);

assert.equal(deduped[0].trackId, "t1");
assert.equal(deduped[0].title, "Song 1");
assert.equal(deduped[1].trackId, "t2");
assert.equal(deduped[1].title, "Song 2");

```

`enforceArtistQuota` という名前の関数を作成してください。

```js
assert.isFunction(enforceArtistQuota)
```

`enforceArtistQuota` は、アーティストが `maxPerArtist` 回を超えて登場しないようにし、最も早いものを保持して余分なトラックを削除する必要があります。

```js
const quotaTracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 121, source: [0, 1], score: 9 },
  { trackId: "t3", artist: "B", title: "Song 3", votes: 1, bpm: 118, source: [0, 2], score: 8 },
  { trackId: "t4", artist: "A", title: "Song 4", votes: 1, bpm: 110, source: [1, 0], score: 0 }
];

const limited = enforceArtistQuota(quotaTracks, 2);

let artistACount = 0;

for (const track of limited) {
  if (track.artist === "A") {
    artistACount++;
  }
}
assert.isAtMost(artistACount, 2);
assert.notInclude(limited.map(t => t.trackId), "t4");
```

`buildSchedule` という名前の関数を作成してください。

```js
assert.isFunction(buildSchedule);
```

`buildSchedule` は `{ slot, trackId }` の形状を持つオブジェクトの配列を返し、`slot` は `1` から始まります。

```js
const scheduleInput = [{ trackId: "t1" }, { trackId: "t2" }, { trackId: "t3" }];
const schedule = buildSchedule(scheduleInput);

assert.lengthOf(schedule, 3);

schedule.forEach((item, index) => {
  assert.hasAllKeys(item, ["slot", "trackId"]);
  assert.equal(item.slot, index + 1);
  assert.equal(item.trackId, scheduleInput[index].trackId);
});
```

`remixPlaylist` という名前の関数を作成してください。

```js
assert.isFunction(remixPlaylist)
```

`remixPlaylist` は最終スケジュールを作成するためにヘルパー関数を順に呼び出す必要があります。

```js
assert.lengthOf(remixPlaylist, 2);

const testPlaylists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t1", artist: "A", title: "Duplicate", votes: 9, bpm: 140 }
  ],
  [
    { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 110 },
    { trackId: "t3", artist: "B", title: "Song 3", votes: 4, bpm: 118 }
  ]
];

const expected = buildSchedule(enforceArtistQuota(dedupeTracks(scoreTracks(flattenPlaylists(testPlaylists))),1));

const actual = remixPlaylist(testPlaylists, 1);

assert.deepEqual(actual, expected);
```

# --seed--

## --seed-contents--

```js
const playlists = [
  [
    {
      trackId: "trk101",
      artist: "Velvet Comet",
      title: "Crimson Afterglow",
      votes: 5,
      bpm: 122
    },
    {
      trackId: "trk102",
      artist: "Neon Harbor",
      title: "Static Horizon",
      votes: 2,
      bpm: 108
    },
    {
      trackId: "trk103",
      artist: "Lunar Arcade",
      title: "Midnight Frequency",
      votes: 4,
      bpm: 128
    }
  ],
  [
    {
      trackId: "trk201",
      artist: "Solar Echo",
      title: "Glass Skyline",
      votes: 3,
      bpm: 115
    },
    {
      trackId: "trk202",
      artist: "Velvet Comet",
      title: "Satellite Hearts",
      votes: 6,
      bpm: 124
    }
  ]
];
```

# --solutions--

```js
function flattenPlaylists(playlists) {
  if (!Array.isArray(playlists)) return [];

  const result = [];

  for (let i = 0; i < playlists.length; i++) {
    const playlist = playlists[i];
    if (!Array.isArray(playlist)) continue;

    for (let j = 0; j < playlist.length; j++) {
      const track = playlist[j];
      result.push({
        trackId: track.trackId,
        artist: track.artist,
        title: track.title,
        votes: track.votes,
        bpm: track.bpm,
        source: [i, j]
      });
    }
  }
  return result;
}

function scoreTracks(tracks) {
  const targetBpm = 120;
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const track = tracks[i];

    result.push({
      trackId: track.trackId,
      artist: track.artist,
      title: track.title,
      votes: track.votes,
      bpm: track.bpm,
      source: track.source,
      score: (track.votes * 10) - Math.abs(track.bpm - targetBpm)
    });
  }
  return result;
}

function dedupeTracks(tracks) {
  const seenTrackIds = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    if (!seenTrackIds.includes(tracks[i].trackId)) {
      seenTrackIds.push(tracks[i].trackId);
      result.push(tracks[i]);
    }
  }
  return result;
}

function enforceArtistQuota(tracks, maxPerArtist) {
  const artists = [];
  const artistCounts = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const artist = tracks[i].artist;
    const artistIndex = artists.indexOf(artist);

    if (artistIndex === -1) {
      artists.push(artist);
      artistCounts.push(1);
      result.push(tracks[i]);
    } else if (artistCounts[artistIndex] < maxPerArtist) {
      artistCounts[artistIndex]++;
      result.push(tracks[i]);
    }
  }

  return result;
}

function buildSchedule(tracks) {
  const schedule = [];

  for (let i = 0; i < tracks.length; i++) {
    schedule.push({
      slot: i + 1,
      trackId: tracks[i].trackId
    });
  }

  return schedule;
}

function remixPlaylist(playlists, maxPerArtist) {
  const flattened = flattenPlaylists(playlists);
  const scored = scoreTracks(flattened);
  const deduped = dedupeTracks(scored);
  const limited = enforceArtistQuota(deduped, maxPerArtist);
  return buildSchedule(limited);
}
```
