---
id: 66a9bb8578459a175432e7d0
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

在表格的第一部分，你會想要將表頭內容群組，該內容表述學生名字、姓氏和期末考成績的行標籤。

<dfn>表格標頭</dfn>元素 `thead` 用於將表格中的標頭內容分組。 

以下是使用 `thead` 元素的範例：

```html
<table>
  <thead>
    <!-- header content goes here -->
  </thead>
</table>
```

在你的 `caption` 元素下方，添加一個表格標頭元素。

# --hints--

你應該有一個開啟的 `thead` 頁籤。

```js
assert.match(code, /<thead>/i);
```

你應該有一個結束的 `thead` 頁籤。

```js
assert.match(code, /<\/thead>/i);
```

你的 `thead` 元素應該在你的 `table` 元素內。

```js
const theadEl = document.querySelector('thead');
assert.strictEqual(theadEl?.parentElement?.tagName, 'TABLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
