---
id: 66a9ca099e3ecb1be10a2696
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

現在是時候為表格添加另外 2 名學生。

遵循前面步驟的相同樣式，為學生表格添加第四個列。使用以下資料作為表格資料元素：

- 姓氏：`Thompson`
- 名字：`Jane`
- 成績：`77`

對於第五個學生表格列，請使用下列資料作為表格資料元素：

- 姓氏：`Williams`
- 名字：`Natalie`
- 成績：`83`

# --hints--

你應該在你的 `tbody` 元素內有第四個 `tr`。

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(4)'));
```

你應該在第四個 `tr` 元素內有一個包含文字 `Thompson` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:first-child')?.textContent.trim(), 'Thompson');
```

你應該在第四個 `tr` 元素內有一個包含文字 `Jane` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(2)')?.textContent.trim(), 'Jane');
```

你應該在第四個 `tr` 元素內有一個包含文字 `77` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(3)')?.textContent.trim(), '77');
```

你應該在你的 `tbody` 元素內有第五個 `tr`。

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 5);
```

你應該在第五個 `tr` 元素內有一個包含文字 `Williams` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:first-child')?.textContent.trim(), 'Williams');
```

你應該在第五個 `tr` 元素內有一個包含文字 `Natalie` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(2)')?.textContent.trim(), 'Natalie');
```

你應該在第五個 `tr` 元素內有一個包含文字 `83` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(3)')?.textContent.trim(), '83');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
