---
id: 66a9c006e4f58819396605b3
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

現在你已經完成了標題部分，是時候添加<dfn>表體</dfn> `tbody` 了。 表體將顯示所有學生姓名及其成績。

爲表格添加表體元素。

# --hints--

你應該有一個 `tbody` 開始標籤。

```js
assert.match(code, /<tbody>/i);
```

你應該有一個 `tbody` 結束標籤。

```js
assert.match(code, /<\/tbody>/i);
```

你的 `tbody` 開始標籤應在 `table` 結束標籤之前。

```js
assert.match(code, /<tbody>[\s\S]*<\/tbody>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

--fcc-editable-region--

--fcc-editable-region--
    </table>
  </body>
</html>
```
