---
id: 66a9ca099e3ecb1be10a2696
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

現在是時候給表格再增加兩名學生了。

按照與前面步驟相同的模式，添加第四個學生表格行。 表格數據元素使用以下數據：

- 姓氏：`Thompson`
- 名字：`Jane`
- 分數：`77`

對於第五個學生表格行，表格數據元素使用以下數據：

- 姓氏：`Williams`
- 名字：`Natalie`
- 分數：`83`

# --hints--

你應該在 `tbody` 元素中添加第四個 `tr`。

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(4)'));
```

你應該在第四個 `tr` 元素中包含一個文本爲 `Thompson` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:first-child')?.textContent.trim(), 'Thompson');
```

你應該在第四個 `tr` 元素中包含一個文本爲 `Jane` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(2)')?.textContent.trim(), 'Jane');
```

你應該在第四個 `tr` 元素中包含一個文本爲 `77` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(3)')?.textContent.trim(), '77');
```

你應該在 `tbody` 元素中添加第五個 `tr`。

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 5);
```

你應該在第五個 `tr` 元素中包含一個文本爲 `Williams` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:first-child')?.textContent.trim(), 'Williams');
```

你應該在第五個 `tr` 元素中包含一個文本爲 `Natalie` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(2)')?.textContent.trim(), 'Natalie');
```

你應該在第五個 `tr` 元素中包含一個文本爲 `83` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(3)')?.textContent.trim(), '83');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

--fcc-editable-region--

--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
