---
id: 66a9bcf00f13a418368a272e
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

表格標題元素由一個 <dfn>表格列</dfn> 元素 `tr` 組成，該元素包含 <dfn>表格標題儲存格</dfn> 元素 `th`。

以下是使用 `tr` 和 `th` 元素建立運動表格的範例：

```html
<table>
  <caption>Football Scores</caption>
  <thead>
    <tr>
      <th>Team</th>
      <th>Wins</th>
      <th>Losses</th>
    </tr>
  </thead>
</table>
```

在你的 `thead` 元素內，添加一個 `tr` 元素。

在你的 `tr` 元素內，為你的 `tr` 元素添加三個 `th` 元素。 

第一個 `th` 元素應包含文字 `Last Name`。第二個 `th` 元素應包含文字 `First Name`。第三個 `th` 元素應包含文字 `Grade`。

# --hints--

你應該有一個開啟的 `tr` 頁籤。

```js
assert.match(code, /<tr>/i);
```

你應該有一個結束的 `tr` 頁籤。

```js
assert.match(code, /<\/tr>/i);
```

你的 `tr` 元素應該在你的 `thead` 元素內。 

```js
const trEl = document.querySelector('tr');
assert.strictEqual(trEl?.parentElement?.tagName, 'THEAD');
```

你的第一個 `th` 元素應該有 `Last Name` 的文字。

```js
const thEl = document.querySelectorAll('th')[0];
assert.strictEqual(thEl?.textContent.trim(), 'Last Name');
```

你的第二個 `th` 元素應該有 `First Name` 的文字。

```js
const thEl = document.querySelectorAll('th')[1];
assert.strictEqual(thEl?.textContent.trim(), 'First Name');
```

你的第三個 `th` 元素應該有 `Grade` 的文字。

```js
const thEl = document.querySelectorAll('th')[2];
assert.strictEqual(thEl?.textContent.trim(), 'Grade');
```

你所有的 `th` 元素都應該在你的 `tr` 元素內。

```js
assert.match(code, /<tr>.*<th>.*<\/th>.*<th>.*<\/th>.*<th>.*<\/th>.*<\/tr>/is);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
      --fcc-editable-region--
        
      --fcc-editable-region--
      </thead>
    </table>
  </body>
</html>
```
