---
id: 66a9c14e3b34c719e34bae20
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

要在表格中添加學生數據，需要使用表格行和表格數據元素。

<dfn>表格數據</dfn>元素 `td`，用於創建表格中的單元格。

下面是一個將 `td` 元素用於體育運動員表格的示例：

```html
<tr>
  <td>1</td>
  <td>John Doe</td>
  <td>USA</td>
</tr>
```

在表體元素中，添加一個表格行元素 `tr` 和三個表格數據元素 `td`。

第一個表格數據元素應包含姓氏 `Davis`。

第二個表格數據元素應包含名字 `Alex`。

第三個表格數據元素應包含分數 `54`。

# --hints--

你應該在 `tbody` 元素中包含一個 `tr` 元素。

```js
assert.isNotNull(document.querySelector('tbody tr'));
```

你應該在 `tr` 元素中包含一個文本爲 `Davis` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr td')?.textContent.trim(), 'Davis');
```

你應該在 `tr` 元素中包含一個文本爲 `Alex` 的 `td` 元素。

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[1]?.textContent.trim(), 'Alex');
```

你應該在 `tr` 元素中包含一個文本爲 `54` 的 `td` 元素。

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[2]?.textContent.trim(), '54');
```

你應將三個 `td` 元素都放在 `tr` 元素中。

```js
assert.match(code, /<tr>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<\/tr>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
--fcc-editable-region--

--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
