---
id: 66a9c7a4521d2b1b1ec6dcf0
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

現在是時候為表格添加另外 2 名學生。

遵循前一步驟相同的樣式，為學生表格添加第二個列。使用以下資料作為表格資料元素：

- 姓氏：`Doe`
- 名字：`Samantha`
- 成績：`92`

對於第三個學生表格列，請使用下列資料作為表格資料元素：

- 姓氏：`Rodriguez`
- 名字：`Marcus`
- 成績：`88`

# --hints--

你應該在你的 `tbody` 元素內有第二個 `tr`。

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(2)'));
```

你應該在第二個 `tr` 元素內有一個包含文字 `Doe` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:first-child')?.textContent.trim(), 'Doe');
```

你應該在第二個 `tr` 元素內有一個包含文字 `Samantha` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:nth-child(2)')?.textContent.trim(), 'Samantha');
```

你應該在第二個 `tr` 元素內有一個包含文字 `92` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:last-child')?.textContent.trim(), '92');
```

你應該在你的 `tbody` 元素內有第三個 `tr`。

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 3);
```

你應該在第三個 `tr` 元素內有一個包含文字 `Rodriguez` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:first-child')?.textContent.trim(), 'Rodriguez');
```

你應該在第三個 `tr` 元素內有一個包含文字 `Marcus` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:nth-child(2)')?.textContent.trim(), 'Marcus');
```

你應該在第三個 `tr` 元素內有一個包含文字 `88` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:last-child')?.textContent.trim(), '88');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
