---
id: 66a9bcf00f13a418368a272e
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

表頭元素由<dfn>表格行</dfn>元素 `tr` 組成，其中包含<dfn>表頭單元格</dfn>元素 `th`。

下面是一個使用 `tr` 和 `th` 元素的運動表格的示例：

```html
<table>
  <caption>Football Scores</caption>
  <thead>
    <tr>
      <th>Team</th>
      <th>Wins</th>
      <th>Losses</th>
    </tr>
  </thead>
</table>
```

在 `thead` 元素中，添加一個 `tr` 元素。

在你的 `tr` 元素中，添加三個 `th` 元素。

第一個 `th` 元素應包含文本 `Last Name`。 第二個 `th` 元素應包含文本 `First Name`。 第三個 `th` 元素應包含文本 `Grade`。

# --hints--

你應該有一個 `tr` 開始標籤。

```js
assert.match(code, /<tr>/i);
```

你應該有一個 `tr` 結束標籤。

```js
assert.match(code, /<\/tr>/i);
```

你的 `tr` 元素應該在 `thead` 元素中。

```js
const trEl = document.querySelector('tr');
assert.strictEqual(trEl?.parentElement?.tagName, 'THEAD');
```

你的第一個 `th` 元素的文本應爲 `Last Name`。

```js
const thEl = document.querySelectorAll('th')[0];
assert.strictEqual(thEl?.textContent.trim(), 'Last Name');
```

你的第二個 `th` 元素的文本應爲 `First Name`。

```js
const thEl = document.querySelectorAll('th')[1];
assert.strictEqual(thEl?.textContent.trim(), 'First Name');
```

你的第三個 `th` 元素的文本應爲 `Grade`。

```js
const thEl = document.querySelectorAll('th')[2];
assert.strictEqual(thEl?.textContent.trim(), 'Grade');
```

你所有的 `th` 元素都應該在 `tr` 元素中。

```js
assert.match(code, /<tr>.*<th>.*<\/th>.*<th>.*<\/th>.*<th>.*<\/th>.*<\/tr>/is);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
      --fcc-editable-region--

      --fcc-editable-region--
      </thead>
    </table>
  </body>
</html>
```
