---
id: 66a9bb8578459a175432e7d0
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

對於表格的第一部分，你需要對標題內容進行分組，標題內容代表學生的名、姓和期末考試成績的列標籤。

<dfn>表頭</dfn>元素 `thead`，用於在表格中對標題內容進行分組。

下面是一個使用 `thead` 元素的示例：

```html
<table>
  <thead>
    <!-- header content goes here -->
  </thead>
</table>
```

在 `caption` 元素下方，添加一個表頭元素。

# --hints--

你應該有一個 `thead` 開始標籤。

```js
assert.match(code, /<thead>/i);
```

你應該有一個 `thead` 結束標籤。

```js
assert.match(code, /<\/thead>/i);
```

你的 `thead` 元素應該在 `table` 元素中。

```js
const theadEl = document.querySelector('thead');
assert.strictEqual(theadEl?.parentElement?.tagName, 'TABLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>
--fcc-editable-region--

--fcc-editable-region--
    </table>
  </body>
</html>
```
