---
id: 66a9c7a4521d2b1b1ec6dcf0
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

現在是時候給表格再增加兩名學生了。

按照與上一步相同的模式，添加第二行學生表。 表格數據元素使用以下數據：

- 姓氏：`Doe`
- 名字：`Samantha`
- 分數：`92`

對於第三個學生表格行，表格數據元素使用以下數據：

- 姓氏：`Rodriguez`
- 名字：`Marcus`
- 分數：`88`

# --hints--

你應該在 `tbody` 元素中添加第二個 `tr`。

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(2)'));
```

你應該在第二個 `tr` 元素中包含一個文本爲 `Doe` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:first-child')?.textContent.trim(), 'Doe');
```

你應該在第二個 `tr` 元素中包含一個文本爲 `Samantha` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:nth-child(2)')?.textContent.trim(), 'Samantha');
```

你應該在第二個 `tr` 元素中包含一個文本爲 `92` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:last-child')?.textContent.trim(), '92');
```

你應該在 `tbody` 元素中添加第三個 `tr`。

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 3);
```

你應該在第三個 `tr` 元素中包含一個文本爲 `Rodriguez` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:first-child')?.textContent.trim(), 'Rodriguez');
```

你應該在第三個 `tr` 元素中包含一個文本爲 `Marcus` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:nth-child(2)')?.textContent.trim(), 'Marcus');
```

你應該在第三個 `tr` 元素中包含一個文本爲 `88` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:last-child')?.textContent.trim(), '88');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

--fcc-editable-region--

--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
