---
id: 66a98f42c7c06903e5f8dd07
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本工作坊，你將通過爲一組學生製作期末考試表格來練習使用 HTML 表格。

開始項目，添加 `<!DOCTYPE html>`，添加一個 `html` 元素，其 `lang` 屬性值爲 `en`。

在 `html` 元素中，添加一個 `head` 元素。

# --hints--

你應該有 `<!DOCTYPE html>`。

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

你應該有一個 `html` 開始標籤，語言設置爲英語。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

你應該有一個 `html` 結束標籤。

```js
assert.match(code, /<\/html>/i);
```

你的 `DOCTYPE` 應放在 `html` 元素之前。

```js
assert.match(code, /<!DOCTYPE\s+html>[.\n\s]*<html\s+lang\s*=\s*('|")en\1\s*>/im)
```

你應該有一個 `head` 開始標籤。

```js
assert.match(code, /<head>/i);
```

你應該有一個 `head` 結束標籤。

```js
assert.match(code, /<\/head>/i);
```

你的 `head` 開始標籤應在 `head` 結束標籤之前。

```js
assert.match(code, /<head>[.\n\s]*<\/head>/im)
```

你的 `head` 元素應該在 `html` 元素中。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[.\n\s]*<head>[.\n\s]*<\/head>[.\n\s]*<\/html>/im);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
