---
id: 66a9c14e3b34c719e34bae20
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

要為表格添加學生資料，你需要使用表格列和表格資料元素。

<dfn>表格資料</dfn>元素 `td` 用於創建表格中的儲存格。

以下是使用 `td` 元素建立運動員表格的範例：

```html
<tr>
  <td>1</td>
  <td>John Doe</td>
  <td>USA</td>
</tr>
```

在你的表格主體元素中，為你的表格添加一個 `tr` 表格列元素，該列包含三個 `td` 表格資料元素。 

第一個表格資料元素應包含 `Davis` 的姓氏。 

第二個表格資料元素應包含 `Alex` 的名字。 

第三個表格資料元素應該包含 `54` 的等級。

# --hints--

你應該在你的 `tbody` 元素內有一個 `tr` 元素。

```js
assert.isNotNull(document.querySelector('tbody tr'));
```

你應該在 `tr` 元素內有一個包含文字 `Davis` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr td')?.textContent.trim(), 'Davis');
```

你應該在 `tr` 元素內有一個包含文字 `Alex` 的 `td` 元素。

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[1]?.textContent.trim(), 'Alex');
```

你應該在 `tr` 元素內有一個包含文字 `54` 的 `td` 元素。

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[2]?.textContent.trim(), '54');
```

你應該在 `tr` 元素內擁有全部三個 `td` 元素。

```js
assert.match(code, /<tr>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<\/tr>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
