---
id: 66a9b8f14b963916a3baa732
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

若要為 `table` 添加標題，可以使用 <dfn>table caption</dfn> 元素。

以下是使用 `caption` 元素的範例：

```html
<table>
  <caption>Football Scores</caption>
</table>
```

在你的 `table` 元素內，巢狀一個 `caption` 元素，並帶有文字 `Calculus Final Exam Grades`。

# --hints--

你應該有一個開啟的 `caption` 頁籤。

```js
assert.match(code, /<caption>/i);
```

你應該有一個結束的 `caption` 頁籤。

```js
assert.match(code, /<\/caption>/i);
```

你的 `caption` 元素應該在你的 `table` 元素內。

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.parentElement.tagName, 'TABLE');
```

你的 `caption` 應該包含文字 `Calculus Final Exam Grades`。

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.innerText.trim(), 'Calculus Final Exam Grades');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
    --fcc-editable-region--
      
    --fcc-editable-region--      
    </table>
  </body>
</html>
```
