---
id: 66a9cb26ec6bd41cf6c82bc5
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

最後要添加到表格中的部分是 <dfn>table foot</dfn> 元素 `tfoot`。表格腳元素將用於顯示表格中所有學生的平均成績。

為你的表格添加 `tfoot` 元素。 

在 `tfoot` 元素內，添加一個 `tr` 元素。

在 `tr` 元素內，添加兩個 `td` 元素。 

第一個 `td` 元素應該包含文字 `Average Grade`。第二個 `td` 元素應該包含文字 `78.8`。

# --hints--

你應該有一個開啟的 `tfoot` 頁籤。

```js
assert.match(code, /<tfoot>/);
```

你應該有一個結束的 `tfoot` 頁籤。

```js
assert.match(code, /<\/tfoot>/);
```

你的 `tfoot` 元素內應該有一個 `tr` 元素。

```js
assert.isNotNull(document.querySelector('tfoot tr'));
```

你的第一個 `td` 元素應該包含文字 `Average Grade`。請確定它位於 `tr` 元素內。

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child')?.textContent.trim(), 'Average Grade');
```

你的第二個 `td` 元素應該包含文字 `78.8`。請確保它位於 `tr` 元素內。

```js
assert.strictEqual(document.querySelector('tfoot tr td:nth-of-type(2)')?.textContent.trim(), '78.8');
```

你的 `tr` 元素內應該有兩個 `td` 元素。

```js
assert.strictEqual(document.querySelectorAll('tfoot tr td').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
