---
id: 66a9cb26ec6bd41cf6c82bc5
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

最後要添加到表格中的部分是<dfn>表腳</dfn>元素 `tfoot`。 表腳元素將用於顯示錶格中所有學生的平均成績。

在表格中添加 `tfoot` 元素。

在 `tfoot` 元素中，添加一個 `tr` 元素。

在 `tr` 元素中，添加兩個 `td` 元素。

第一個 `td` 元素應包含文本 `Average Grade`。 第二個 `td` 元素應包含文本 `78.8`。

# --hints--

你應該有一個 `tfoot` 開始標籤。

```js
assert.match(code, /<tfoot>/);
```

你應該有一個 `tfoot` 結束標籤。

```js
assert.match(code, /<\/tfoot>/);
```

你應該在 `tfoot` 元素中包含一個 `tr` 元素。

```js
assert.isNotNull(document.querySelector('tfoot tr'));
```

你的第一個 `td` 元素應包含文本 `Average Grade`。 確保它位於 `tr` 元素內。

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child')?.textContent.trim(), 'Average Grade');
```

你的第二個 `td` 元素應包含文本 `78.8`。 確保它位於 `tr` 元素內。

```js
assert.strictEqual(document.querySelector('tfoot tr td:nth-of-type(2)')?.textContent.trim(), '78.8');
```

你的 `tr` 元素應包含兩個 `td` 元素。

```js
assert.strictEqual(document.querySelectorAll('tfoot tr td').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>

--fcc-editable-region--

--fcc-editable-region--
    </table>
  </body>
</html>
```
