---
id: 66a9b8f14b963916a3baa732
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

要爲 `table` 添加標題，你可以使用 <dfn>table caption</dfn> 元素。

下面是一個使用 `caption` 元素的示例：

```html
<table>
  <caption>Football Scores</caption>
</table>
```

在你的 `table` 元素中，嵌套一個 `caption` 元素，文本爲 `Calculus Final Exam Grades`。

# --hints--

你應該有一個 `caption` 開始標籤。

```js
assert.match(code, /<caption>/i);
```

你應該有一個 `caption` 結束標籤。

```js
assert.match(code, /<\/caption>/i);
```

你的 `caption` 元素應該在 `table` 元素中。

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.parentElement.tagName, 'TABLE');
```

你的 `caption` 應包含文本 `Calculus Final Exam Grades`。

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.innerText.trim(), 'Calculus Final Exam Grades');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
    --fcc-editable-region--

    --fcc-editable-region--      
    </table>
  </body>
</html>
```
