---
id: 66a9cd4ffa48cf1ef7333640
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

你的表格幾乎完成，但還有最後一件事要添加。

如果用於 `Average Grade` 的 `td` 元素能跨越兩個行而不只是跨越一個行，那會更好。

為此，你可以在 `td` 元素上使用 <dfn>colspan</dfn> 屬性。 

`colspan` 屬性指定儲存格應該跨越的行數。

以下是使用 `colspan` 屬性於運動表格的範例：

```html
<tr>
  <td colspan="3">Total Points</td>
</tr>
```

在開啟的 `td` 頁籤內，為你的 `colspan` 屬性添加並設定為 `"2"`。

有了這個變更，你的表格就完成了！

# --hints--

你不應該更改現有 `td` 元素的內容。

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child').textContent.trim(), 'Average Grade');
```

你的 `td` 元素應該有一個 `colspan` 屬性。

```js
assert.isTrue(document.querySelector('tfoot tr td:first-child').hasAttribute('colspan'));
```

你的 `td` 元素應該有一個值為 `"2"` 的 `colspan` 屬性。

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child').getAttribute('colspan'), '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>

      <tfoot>
        <tr>
          --fcc-editable-region--
          <td>Average Grade</td>
          --fcc-editable-region--
          <td>78.8</td>
        </tr>
      </tfoot>
    </table>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>

      <tfoot>
        <tr>
          <td colspan="2">Average Grade</td>
          <td>78.8</td>
        </tr>
      </tfoot>
    </table>
  </body>
</html>
```
