---
id: 66a9b48a211a73155621d0df
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

在這個工作坊中，你將透過建置一個學生群組的期末考試表格來練習操作 HTML 表格。 

在之前的課程中，你學會了如何使用 `table` 元素來表述表格資料。

在你的 `body` 元素內巢狀一個 `table` 元素。

# --hints--

你應該有一個開啟的 `table` 頁籤。

```js
assert.match(code, /<table>/i);
```

你應該有一個結束的 `table` 頁籤。

```js
assert.match(code, /<\/table>/i);
```

你的 `table` 元素應該在你的 `body` 元素內。

```js
const tableEl = document.querySelector('table');
assert.strictEqual(tableEl?.parentElement?.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
