---
id: 66a9c006e4f58819396605b3
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

現在你已經完成了 head 區段，是時候為你的 <dfn>table body</dfn>、`tbody` 添加內容。表格主體將表述所有學生姓名及其成績。

為你的表格添加一個表格主體元素。

# --hints--

你應該有一個開啟的 `tbody` 頁籤。

```js
assert.match(code, /<tbody>/i);
```

你應該有一個結束的 `tbody` 頁籤。

```js
assert.match(code, /<\/tbody>/i);
```

你的開啟 `tbody` 頁籤應該在關閉 `table` 頁籤之前。

```js
assert.match(code, /<tbody>[\s\S]*<\/tbody>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
