---
id: 66a9b333487b9c14998539a5
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `head` 元素內，嵌套一個 `meta` 元素，其 `charset` 屬性的值設置爲 `"UTF-8"`。

在該 `meta` 元素下方，添加一個 `title` 元素。

`title` 元素的文本應爲 `Calculus Final Exams Table`。

完成 `head` 元素後，你可以添加 `body` 元素。

# --hints--

你應該有一個 `meta` 元素。

```js
assert.isNotNull(document.querySelector("meta"));
```

`meta` 元素是一個空元素，它不應該有一個結束標籤 `</meta>`。

```js
assert.notMatch(code, /<\/meta>/i);
```

你的 `meta` 標籤應該有一個 `charset` 屬性。

```js
assert.match(code, /<meta\s+charset\s*/i);
```

你的 `charset` 屬性的值應爲 `"UTF-8"`。

```js
assert.match(code, /charset\s*=\s*('|")UTF-8\1/i);
```

你的 `meta` 元素應嵌套在 `head` 元素中。

```js
const meta = document.querySelector('head > meta');
assert.strictEqual(meta?.parentElement?.tagName, 'HEAD');
```

你應該有一個 `title` 開始標籤。

```js
assert.match(code, /<title>/i);
```

你應該有一個 `title` 結束標籤。

```js
assert.match(code, /<\/title>/i);
```

所有的 `title` 元素都應嵌套在 `head` 元素中。

```js
assert.match(code, /<head>.*\s*<title>.*<\/title>.*\s*<\/head>/si);
```

你的 `title` 元素應包含文本 `Calculus Final Exams Table`。 你可能需要檢查一下拼寫。

```js
const titleText = document.querySelector('title')?.innerText.trim();
assert.strictEqual(titleText, "Calculus Final Exams Table");
```

你應該有一個 `body` 開始標籤。

```js
assert.match(code, /<body>/i);
```

你應該有一個 `body` 結束標籤。

```js
assert.match(code, /<\/body>/i);
```

你不應該改變你的 `head` 元素。 確認你沒有刪除結束標籤。

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

你的 `body` 元素應該在你的 `head` 元素之後。

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
  --fcc-editable-region--
</html>
```
