---
id: 66a9b48a211a73155621d0df
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

In previous lessons, you learned how to work with the `table` element to represent tabular data.

在 `body` 元素內，嵌套一個 `table` 元素。

# --hints--

你應該有一個 `table` 開始標籤。

```js
assert.match(code, /<table>/i);
```

你應該有一個 `table` 結束標籤。

```js
assert.match(code, /<\/table>/i);
```

你的 `table` 元素應該在 `body` 元素中。

```js
const tableEl = document.querySelector('table');
assert.strictEqual(tableEl?.parentElement?.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
