---
id: 6140c7e645d8e905819f1dd4
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

为 `./styles.css` 文件添加一个 `link` 元素，来开始。

# --hints--

您的代码中应该有一个 `link` 元素。

```js
assert.match(code, /<link/);
```

你应该有一个 `link` 元素。

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

你的 `link` 元素应该在你的 `head` 元素中。

```js
assert.exists(document.querySelector('head > link'));
```

你的 `link` 元素应具有一个 `rel` 属性，其值为 `stylesheet`。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

你的 `link` 元素应具有一个 `href` 属性，其值为 `styles.css`。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>

  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css

```
