---
id: 61695c4aad56f95497c19583
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

Remember that the `title` element gives search engines extra information about the page. It also displays the content of that `title` element in two more ways:

* in the title bar when the page is open
* 当你把鼠标悬停在该页面上时，在浏览器标签中。 即使该标签未被激活，一旦你将鼠标悬停在该标签上，`title` 文本就会显示出来。

在 `head` 元素中，嵌套一个带有文本 `Colored Markers` 的 `title` 元素。

# --hints--

你应该有一个 `title` 开始标签。

```js
assert(code.match(/<title\s*>/i));
```

你应该有一个 `title` 结束标签。

```js
assert(code.match(/<\/title\s*>/i));
```

你的项目应该有标题 `Colored Markers`。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

请记住，标题文本的大小写和拼写很重要。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
