---
id: 61696ef7ac756c829f9e4048
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

在这个项目中，你将使用外部 CSS 文件来设置页面样式。 我们已经为你创建了一个 `styles.css` 文件。 但在你使用它之前，你需要将它链接到页面上。

在 `head` 元素中嵌套一个 `link` 元素。 给它一个 `rel` 属性，设置为 `"stylesheet"`，和一个 `href` 属性，设置为 `"styles.css"`。

# --hints--

你应该有一个 `link` 元素。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

你的 `link` 元素应该是一个空元素，它不应该有结束标签 `</link>`。

```js
assert.notMatch(code, /<\/link>/i);
```

你的 `link` 元素应该具有值为 `"stylesheet"` 的 `rel` 属性。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

你的 `link` 元素应该具有值为 `"styles.css"` 的 `href` 属性。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

你的 `link` 元素应该在你的 `head` 元素中。

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
