---
id: 61695197ac34f0407e339882
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

As you've seen in the previous projects, webpages should start with a `DOCTYPE html` declaration, followed by an `html` element.

在文档顶部添加 `DOCTYPE html` 声明，然后添加 `html` 元素。 给 `html` 元素一个值为 `en` 的 `lang` 属性。

# --hints--

代码应包含 `DOCTYPE` 声明。

```js
assert(code.match(/<!DOCTYPE/gi));
```

应该在 `DOCTYPE` 声明之后包含一个空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你应该把文档类型定义为 `html`。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

应该在 `DOCTYPE` 声明的类型后面加上一个 `>`。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你的 `html` 元素应该有一个值为 `en` 的 `lang` 属性的开始标签。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 元素应该有一个结束标签。

```js
assert(code.match(/<\/html\s*>/gi));
```

你的 `DOCTYPE` 声明应该位于 HTML 的开始。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
