---
id: 61695d1fbc003856628bf561
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

`charset` 属性指定了文档使用的字符编码。 `utf-8`（Unicode Transforming Format - 8 bit）是一种用于电子通信的字符编码标准。

在 `head` 元素内，嵌套一个 `meta` 元素，其中 `charset` 属性值为 `"utf-8"`。

# --hints--

你应该有一个 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 1);
```

你的 `meta` 元素应该是一个空元素，它不应该有结束标签 `</meta>`。

```js
assert.notMatch(code, /<\/meta>/i);
```

你的 `meta` 元素应该有一个 `charset` 属性，其值为 `"utf-8"`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

你的 `meta` 元素应该在 `head` 元素中。

```js
const metaElementRegex = /<head\s*>(?:.|\r|\n)*?<meta(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, metaElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
