---
id: 616d595270d902f0e7086e18
title: 步骤15
challengeType: 0
dashedName: step-15
---

# --description--

While you have three separate marker `div` elements, they look like one big rectangle. You should add some space between them to make it easier to see each element.

当简写 `margin` 属性有两个值时，它将 `margin-top` 和 `margin-bottom` 设置为第一个值，并将 `margin-left` 和 `margin-right` 设置为第二个值。

在你的 `.marker` CSS 规则中，将 `margin` 属性设置为 `10px auto`。

# --hints--

你的 `.marker` CSS 规则应该有一个 `margin` 属性设置为 `10px auto`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}
--fcc-editable-region--

```
