---
id: 64dc8b1faf9c60304ca8a9b6
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このプロジェクトでは、ハノイの塔として知られる数学パズルを解きます。このパズルは3本の棒と、直径の異なる複数の円盤で構成されています。

このパズルの目的は、特定のルールに従って、より大きな円盤をより小さな円盤の上に置くことを禁止しながら、円盤を最初の棒から3番目の棒に移動させることです。

まず、棒を表す空の辞書を`rods`という名前で作成してください。

# --hints--

`rods`という名前の変数を用意してください。

```js
({ test: () => assert(__userGlobals.has('rods')) })
```

`rods`変数は空の辞書である必要があります。

```js
({ test: () => assert(runPython(`
    rods == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
