---
id: 64dcd3d61c448e2676501f43
title: ステップ 20
challengeType: 20
dashedName: step-20
---

# --description--

移動番号を3で割った余りが2の場合、移動は`'A'`と`'B'`（元の棒と補助の棒）の間で許可されます。

そのための`elif`文を追加してください。条件が満たされた場合、適切な文字列を出力してください。

# --hints--

`elif`のときに実行する`remainder == 2`文が必要です。

```js
({test: () => assert.match(code, /elif\s+remainder\s*==\s*2:/)})
```

文字列`f'Move {i + 1} allowed between {source} and {auxiliary}'`を出力する必要があります。

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*source\s*\}\sand\s\{\s*auxiliary\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
