---
id: 64dcdee07a14f0299ec30559
title: ステップ 22
challengeType: 20
dashedName: step-22
---

# --description--

ロッド間の許可された移動を見つけるコードを書きましたが、実際の移動は両方向に行われる可能性があります。
`print()` 呼び出しの後に、`forward` という名前の変数を宣言し、`False` に設定してください。その変数を使って、ロッド間でディスクをどの方向に移動させる必要があるかを確認します。

# --hints--

`forward` という名前の変数を宣言してください。

```js
({ test: () => assert.match(code, /forward\s*=\s*./) })
```

`forward` 変数は `False` に設定する必要があります。

```js
({ test: () => assert.match(code, /forward\s*=\s*False/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
--fcc-editable-region--
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
--fcc-editable-region--
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
