---
id: 64dc985e6720a23edac63b51
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

ロッド間で許可されているディスクの移動は、3回ごとに繰り返されるパターンを示します。たとえば、ロッド `A` とロッド `C` 間の移動は、移動番号を3で割った余りが1になる最初、4回目、7回目の移動で許可されます。

以前作成した `for` ループの中で、既存の `print()` 呼び出しを、`if` のときに実行される `(i + 1) % 3 == 1` 文に置き換えてください。この `if` 文の中で、f文字列を使って `f'Move {i + 1} allowed between {source} and {target}'` を出力してください。なお、`i + 1` は移動番号であり、最初の繰り返しでは `i` がゼロであることに注意してください。

# --hints--

`if` のときに実行される `(i + 1) % 3 == 1` 文があるはずです。

```js
({test: () => assert.match(code, /if\s+\(\s*i\s*\+\s*1\s*\)\s*%\s*3\s*==\s*1\s*:/)})
```

f文字列を使って文字列 `f'Move {i + 1} allowed between {source} and {target}'` を出力してください。

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*source\s*\}\sand\s\{\s*target\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        print(i)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
