---
id: 64dc97005bc7943e2978df0a
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

`print()` の呼び出しを削除してください。それから、`move()` 関数の中で `pass` キーワードを削除し、`rods` 辞書の内容を出力してください。

# --hints--

コード内に `print(number_of_moves)` があってはいけません。

```js
({ test: () => assert.isFalse(/print\s*\(\s*number_of_moves\s*\)/.test(code)) })
```

`pass` キーワードを削除し、`rods` 辞書を出力してください。

```js
({ test: () => assert.match(code, /def\s+move\s*\(\s*\)\s*:\s+print\s*\(\s*rods\s*\)(?!\s*pass)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
print(number_of_moves)
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
--fcc-editable-region--
```
