---
id: 64dc976bf864693e668d67e8
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

関数を呼び出して、ターミナルに出力を表示してください。

# --hints--

`move()` 関数を呼び出す必要があります。

```js
({test: () => assert.match(code, /^move\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    print(rods)
--fcc-editable-region--

--fcc-editable-region--
```
