---
id: 64dc916138dbe03bdfe6cec6
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

ハノイの塔の目標は、すべての円盤を最後の棒に移動することです。そのためには、次の3つの簡単なルールに従う必要があります。

1. 一番上の円盤だけを動かせます
2. 一度に動かせる円盤は1つだけです
3. 大きい円盤を小さい円盤の上に置くことはできません

既存のコードの下に、`move` という名前の空の関数を宣言してください。後でその関数を使って円盤を棒の間で動かします。今はエラーを避けるために、関数の中身に `pass` キーワードを使ってください。

# --hints--

`move` という名前の空の関数を宣言してください。関数の中身には正しいインデントで `pass` キーワードを使うことを忘れないでください。

```js
({ test: () => {
		assert(runPython(`
        import inspect
        inspect.isfunction(move)
    `))
	} 
})
```

# --seed--

## --seed-contents--

```py
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--
```
