---
id: 64dc92a9718fb53ca3d1bad7
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

ハノイの塔のパズルは、ディスクの数を n とすると、2<sup>n</sup> - 1 手で解くことができます。`number_of_moves` という名前の変数を宣言し、この変数に合計の手数を代入してください。

Pythonでのべき乗演算子は `**` です。

# --hints--

`number_of_moves` という名前の変数を宣言してください。

```js
({ test: () => assert(__userGlobals.has('number_of_moves')) })
```

`number_of_moves` の値は、手数を計算する式にしてください。

```js
({ test: () => {
    assert(runPython(`  
    number_of_moves == 2**NUMBER_OF_DISKS -1      
    `))
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
```
