---
id: 64dc922df2919e3c38cead61
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

コードの先頭で、ディスクの数を格納するための `NUMBER_OF_DISKS` という名前の変数を宣言し、値を `3` に設定してください。それから、`range()` 関数に渡される最初の引数を新しい変数に置き換えてください。

# --hints--

`NUMBER_OF_DISKS` という名前の変数を宣言してください。

```js
({test: () => assert(__userGlobals.has('NUMBER_OF_DISKS')) })
```

`NUMBER_OF_DISKS` 変数は `3` の値を持つべきです。

```js
({test: () => assert.equal(__userGlobals.get('NUMBER_OF_DISKS'), 3) })
```

`range()` 関数の最初の引数の値を `NUMBER_OF_DISKS` 変数に変更してください。

```js
({test: () => assert.match(code, /('|")A\1\s*:\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

def move():
    pass
```
