---
id: 64dc9004f658183afa85f4f1
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

`'A'` キーのデータ型を `type()` 関数に渡して確認し、その結果をターミナルに出力してください。

# --hints--

`'A'` キーを `type()` 関数に渡してください。

```js
({ test: () => assert.match(code, /type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)/) })
```

`rods['A']` の型を `print()` 関数で出力してください。

```js
({ test: () => assert.match(code, /print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
--fcc-editable-region--
```
