---
id: 64dc8ea01436383a88256d7f
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

パズルは、最初のロッドにディスクが積み重なっている状態から始まります。ディスクはサイズが小さいものが上、大きいものが下になるように並んでいます。最初のリストに、さまざまなディスクのサイズを表す数字を入れてください。

最初のリストにアイテムを手動で追加する代わりに、`3` 関数を使って `1` から `range()` までの数字をカウントダウンするシーケンスを生成し、`rods['A']` に代入してください。ここで、`3` は積み重ねの一番下にある最大のディスクを表し、`1` は一番上にある最小のディスクを表します。

構文は `range(x, y, h)` で、`x` が開始整数（含む）、`y` が終了整数（含まない）、`h` がシーケンス内の隣接する数字の差です。

# --hints--

`range()` 関数を使って、`rods['A']` に数字のシーケンスを代入してください。`range()` 関数の呼び出し構文は `range(x, y, h)` です。

```js
({ test: () => assert(runPython(`
  type(rods['A']) is range
  `))
})
```

`range()` 関数を使って、`3` から `1` までの数字のシーケンスを `rods['A']` に代入してください。

```js
({ test: () => assert(runPython(`
 rods['A'] == range(3, 0, -1)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': [],
    'B': [],
    'C': []
}
--fcc-editable-region--
```
