---
id: 64dc90949ff85e3b37be40d0
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

`range()` 関数は変更不可能な数値のシーケンスを返します。ご覧の通り、`rods['A']` のデータ型は `range` ですが、リストにしたいです。

それを実現するには、`range()` の呼び出しを `list()` 関数に渡してください。

# --hints--

`range()` の呼び出しを `list()` 関数に渡す必要があります。

```js
({ test: () => assert(runPython(`
  rods['A'] == list(range(3, 0, -1))
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
