---
id: 64de1469fccce61940f498c4
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

`rods` 辞書は、3本の棒とそれぞれのディスクを表します。キーとして文字列 `'A'`、`'B'`、`'C'` を与え、それぞれを空のリストに設定してください。

# --hints--

`rods` 辞書には `'A'` キーが必要です。

```js
({ test: () => assert(runPython(`'A' in rods`)) })
```

`rods['A']` は空のリストであるべきです。

```js
({ test: () => assert(runPython(`
    rods['A'] == []
  `))
})
```

`rods` 辞書には `'B'` キーが必要です。

```js
({ test: () => assert(runPython(`'B' in rods`)) })
```

`rods['B']` は空のリストであるべきです。

```js
({ test: () => assert(runPython(`
    rods['B'] == []
  `))
})
```

`rods` 辞書には `'C'` キーが必要です。

```js
({ test: () => assert(runPython(`'C' in rods`)) })
```

`rods['C']` は空のリストであるべきです。

```js
({ test: () => assert(runPython(`
    rods['C'] == []
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {}
--fcc-editable-region--
```
