---
id: 64df3f1011888113fbd3d81b
title: ステップ 44
challengeType: 20
dashedName: step-44
---

# --description--

再帰でパズルを解くには、まず元の問題をより小さなサブ問題に分解することが必要です。

`n`枚のディスクを3番目の棒に減少順に積み上げた最終構成は、次の動きを行うことで得られます。

- `n - 1`枚のディスクを元の棒から補助の棒へ移動する
- 一番大きいディスクを元の棒から目的の棒へ移動する
- そして`n - 1`枚のディスクを補助の棒から目的の棒へ移動する

したがって、`move`関数が最初に行うべきことは、`n - 1`を最初の引数として自分自身を呼び出すことです。しかし、基本ケースを定義せずにこれを行うと、`RecursionError`が発生します。これは関数が無限に自分自身を呼び出し続けるためです。

コメントと`print()`の呼び出しの前に、`n - 1`を最初の引数として再帰関数の呼び出しを追加し、関数本体が`n`がゼロより大きい場合にのみ実行されるようにしてください。今のところ、他の引数は同じ順序のままにしてください。

# --hints--

`move`関数の本体は、`if`がゼロより大きい場合に実行される`n`文で始まるべきです。

```js
({ test: () => assert.match(code, /def\s+move\(\s*n\s*,\s*source\s*,\s*auxiliary\s*,\s*target\s*\)\s*:\s+if\s+n\s*>\s*0\s*:/) })
```

`move(n - 1, source, auxiliary, target)`文の中で`if`を呼び出す必要があります。`print()`の呼び出しはインデントしてください。

```js
const transformedCode = "\n" + code.replace(/\r/g, '');
const move = __helpers.python.getDef(transformedCode, "move");
const { function_parameters, function_body } = move;

assert.match(function_parameters, /n\s*,\s*source\s*,\s*auxiliary\s*,\s*target/);
const functionBodyIndent = function_body.match(/ +/)[0];
const re = new RegExp(`${functionBodyIndent}if\\s+n\\s*>\\s*0\\s*:\\s+^${functionBodyIndent}( +)move\\(\\s*n\\s*-\\s*1\\s*,\\s*source\\s*,\\s*auxiliary\\s*,\\s*target\\s*\\).+?^${functionBodyIndent}\\1print\\s*\\(\\s*rods\\s*,\\s*("|')\\\\n\\2\\s*\\)`, "ms");
assert.match(function_body, re);
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods, '\n')
    
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
