---
id: 64df47b32b92301a815d5ef8
title: ステップ 49
challengeType: 20
dashedName: step-49
---

# --description--

最初に、追加した再帰呼び出しは `n - 1` 個の円盤を2番目の棒に移動するというサブ問題を扱っています。

そのため、`target` の引数は2番目の棒に対応し、`auxiliary` の引数は3番目の棒に対応します。これらは再帰が進むにつれて入れ替わり続けることを覚えておいてください。

再帰呼び出しの中で `target` と `auxiliary` の引数の順序を入れ替えて修正してください。

# --hints--

`move(n - 1, source, auxiliary, target)` の呼び出しで引数の順序を変更する必要があります。

```js
({ test:  () => assert.match(code, /move\(\s*n\s*-\s*1\s*,\s*source\s*,\s*target\s*,\s*auxiliary\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, auxiliary, target)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
