---
id: 64df496c6a8ddf1b38db1ed6
title: ステップ 50
challengeType: 20
dashedName: step-50
---

# --description--

前のステップでは、部分問題の中で一番大きなディスクをターゲットの棒に移動するコードを書きました。

今やるべきことは、すでに移動させた`n - 1`枚のディスクを移動させるために、もう一つの再帰呼び出しを追加することだけです。最初の再帰呼び出しをコピーして、`if`ブロックの最後に貼り付けてください。

関数の引数の順番が正しくないことに注意してください。正しい順番を見つけてみましょう。

# --hints--

`move(n - 1,  auxiliary, source, target)`関数の最後に`move`を呼び出す必要があります。

```js
({ test: () => assert.match(code, /^\s{8}move\(\s*n\s*-\s*1\s*,\s*auxiliary\s*,\s*source\s*,\s*target\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
