---
id: 650c6082e5586f9e3acfcd3b
title: ステップ 52
challengeType: 20
dashedName: step-52
---

# --description--

今、`rods` 辞書を削除し、そのキーをそれぞれ変数 `A`、`B`、`C` にして値はそのままにしてください。これらの変更を反映するようにコードをリファクタリングしてください。ターミナルに出力が表示されれば、このステップを正しく行えています。

# --hints--

`rods['A']` を変数 `A` にして値はそのままにしてください。

```js
assert.match(code, /A\s*=\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/);
```

`rods['B']` を変数 `B` にして値はそのままにしてください。

```js
assert.match(code, /B\s*=\s*\[\s*\]/);
```

`rods['C']` を変数 `C` にして値はそのままにしてください。

```js
assert.match(code, /C\s*=\s*\[\s*\]/);
```

`print()` の呼び出しを修正して、`A` オブジェクトの代わりに `B`、`C`、`rods` を出力してください。`print()` の呼び出しでは改行文字を保持してください。

```js
assert.match(code, /print\s*\(\s*A\s*,\s*B\s*,\s*C\s*,\s*('|")\\n\1\s*\)/);
```

コード内に `target.append(source.pop())` があるはずです。

```js
assert.match(code, /target\.append\s*\(\s*source\.pop\s*\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
--fcc-editable-region--
```
