---
id: 65782342850feb3b8d62f936
title: ステップ 18
challengeType: 20
dashedName: step-18
---

# --description--

`(i + 1) % 3` 式を何度も使うので、変数にストアしておくと便利です。

`if` 文のすぐ上に、`remainder` 変数を宣言して、この変数に `(i + 1) % 3` の値を代入してください。

# --hints--

`remainder` ブロックのすぐ上に `if` 変数を宣言してください。

```js
({ test: () => assert.match(code, /^\s{8}remainder\s*=/m) })
```

`remainder` 変数の値は `(i + 1) % 3` 式にしてください。

```js
({ test: () => assert.match(code, /^\s{8}remainder\s*=\s*\(\s*i\s*\+\s*1\s*\)\s*%\s*3/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        
        if (i + 1) % 3 == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
