---
id: 6391d1a4f7ac71efd0621380
title: Rezeptseite erstellen
challengeType: 25
dashedName: top-build-a-recipe-project
---

# --description--

Die Website besteht aus einer Hauptindexseite, die Links zu einigen Rezepten enthält. Die Website wird nicht besonders hübsch aussehen, wenn Sie fertig sind.

**User Stories:**

1. Ihre Rezeptseite sollte ein `DOCTYPE`-Tag enthalten.
1. Ihre Rezeptseite sollte ein `html`-Element mit einem `head`- und `body`-Element als Kinder enthalten.
1. Sie sollten ein `title`-Element innerhalb des `head`-Elements mit dem Text `The Odin Recipes` haben.
1. Sie sollten ein `h1`-Element sehen, das den Text `Creamy Chocolate Fudge` enthält.
1. Sie sollten ein passendes Bild mit einem `alt`-Attribut sehen.
1. Unter dem Bild sollte ein `h2`-Element mit dem Text `Description` sein.
1. Unter `Description` sollten Sie ein paar Absätze sehen, die das Rezept beschreiben.
1. Es sollte ein `h2`-Element mit dem Text `Ingredients` geben.
1. Unter der `Ingredients`-Überschrift sollte eine ungeordnete Liste mit den Zutaten für das Rezept sein.
1. Unter der Zutatenliste fügen Sie eine weitere Überschrift namens `Steps` hinzu.
1. Sie sollten eine geordnete Liste mit ein paar Schritten sehen, die nötig sind, um das Rezept fertigzustellen.
1. Unter den Schritten sollte ein `h2`-Element mit dem Text `More Recipes` sein.
1. Sie sollten ein paar Links zu anderen Rezepten in einer ungeordneten Liste sehen, die ein paar Listenelemente mit Ankerelementen enthält.
1. Diese Ankerelemente sollten ein `href`-Attribut mit dem Wert `#` haben.

# --hints--

Sie sollten ein `DOCTYPE`-Tag haben.

```js
assert(code.match(/<!DOCTYPE\s+?html\s*?>/gi));
```

Sie sollten ein `html`-Element mit `head`- und `body`-Element haben.

```js
const html = document.querySelectorAll('html')[0];
const head = document.querySelectorAll('html > head')[0];
const body = document.querySelectorAll('html > body')[0];

assert(html && head && body);
```

Sie sollten ein `title`-Element innerhalb des `head`-Elements haben, das den Text `The Odin Recipes` enthält.

```js
assert(document.querySelectorAll('HEAD > TITLE')[0]?.innerText == 'The Odin Recipes');
```

Sie sollten ein `h1`-Element innerhalb Ihres `body`-Elements haben, das den Text `Creamy Chocolate Fudge` enthält.

```js
assert(document.querySelectorAll('BODY > H1')[0]?.innerText == 'Creamy Chocolate Fudge');
```

Sie sollten ein Bild mit einem `alt`-Attribut haben.

```js
const img = document.querySelectorAll('IMG')[0];

assert(img && img.alt !='' && img.src != '')
```

Sie sollten ein `h2`-Element mit dem Text `Description` haben.

```js
const h2 = document.querySelectorAll('H2')[0];

assert(h2?.innerText == 'Description');
```

Sie sollten mindestens zwei `p`-Elemente haben, die das Rezept beschreiben.

```js
const paragraphs = document.querySelectorAll('P');

assert(paragraphs?.length > 1);
```

Sie sollten ein `h2`-Element mit dem Text `Ingredients` haben.

```js
const h2 = document.querySelectorAll('H2')[1];

assert(h2?.innerText == 'Ingredients');
```

Sie sollten ein `ul`-Element verschachtelt mit `li`-Elementen haben, die die Zutaten enthalten.

```js
const headers = document.querySelectorAll("h2");

headers.forEach(header => {
    if (header.textContent.trim() === "Ingredients") {
        const next = header.nextElementSibling;

        if(next.tagName === "UL") {
            const listItems = next.querySelectorAll("LI");

            assert(listItems.length > 1);
        } else {
            assert(false);
        }
    }
});
```

Sie sollten ein `h2`-Element mit dem Text `Steps` haben.

```js
const h2 = document.querySelectorAll('H2')[2];

assert(h2?.innerText == 'Steps');
```

Sie sollten ein `ol`-Element haben, das `li`-Elemente mit den Schritten des Rezepts enthält.

```js
const orderedList = document.querySelectorAll('OL')[0];
const listItems = document.querySelectorAll('OL > LI');

assert(orderedList && listItems && listItems.length > 1);
```

Sie sollten ein `h2`-Element mit dem Text `More Recipes` haben.

```js
const h2 = document.querySelectorAll('H2')[3];

assert(h2?.innerText == 'More Recipes');
```

Sie sollten ein `ul`-Element verschachtelt mit `li`-Elementen haben, die `a`-Elemente enthalten, die zu anderen Rezepten verlinken.

```js
const headers = document.querySelectorAll("h2");

headers.forEach(header => {
    if (header.textContent.trim() === "More Recipes") {
        const next = header.nextElementSibling;

        if(next.tagName === "UL") {
            const listItems = next.querySelectorAll("LI");

            assert(listItems.length > 1);

            listItems.forEach(listItem => {
                const anchor = listItem.querySelector("A");

                assert(anchor);
            });
        } else {
            assert(false);
        }
    }
});
```

Ihre Ankerelemente, die zu den Rezepten verlinken, sollten ein `href`-Attribut mit dem Wert `#` haben.

```js
const anchorTags = document.querySelectorAll("a");

const allAnchorsHaveHrefHash = [...anchorTags].every(function(anchorTag) {
  return anchorTag.hasAttribute("href") && anchorTag.getAttribute("href") === "#";
});

assert(allAnchorsHaveHrefHash && anchorTags.length > 0); 
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>The Odin Recipes</title>
  </head>
  <body>
    <h1>Creamy Chocolate Fudge</h1>
    <img src="https://cdn.freecodecamp.org/curriculum/odin-project/build-a-recipe-page/build-a-recipe-page-01.jpg" alt="A delicious chocolate fudge dessert">
    <h2>Description</h2>
    <p>This recipe is for a rich and creamy chocolate fudge that is sure to satisfy your sweet tooth. It's perfect for a special occasion or as a tasty treat for any time of the year.</p>
    <p>This recipe is easy to follow and only requires a few simple ingredients. With just a few steps, you'll be able to create a delicious dessert that everyone will love.</p>
    <h2>Ingredients</h2>
    <ul>
      <li>1 cup sugar</li>
      <li>1/2 cup unsalted butter</li>
      <li>1/4 cup milk</li>
      <li>1/4 cup cocoa powder</li>
      <li>1/4 cup chocolate chips</li>
      <li>1/4 tsp salt</li>
      <li>1 tsp vanilla extract</li>
    </ul>
    <h2>Steps</h2>
    <ol>
      <li>In a medium saucepan, melt the butter over medium heat.</li>
      <li>Add the sugar, milk, cocoa powder, and salt to the saucepan and stir until well combined.</li>
      <li>Bring the mixture to a boil, stirring constantly, and then reduce the heat to low and simmer for 5 minutes.</li>
      <li>Remove the saucepan from the heat and stir in the chocolate chips and vanilla extract until the chocolate is melted and the mixture is smooth.</li>
      <li>Pour the fudge into a greased 8-inch square pan and let it cool completely before cutting into squares.</li>
    </ol>
    <h2>More Recipes</h2>
    <ul>
      <li><a href="#">Peanut Butter Cookies</a></li>
      <li><a href="#">Lemon Bars</a></li>
      <li><a href="#">Chocolate Chip Pancakes</a></li>
    </ul>
  </body>
</html>
```

```css

```
