---
id: 642db8c409d9991d0b3b2f0d
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

El proyecto empieza con un contenedor HTML básico y algo de CSS correspondiente. Tu primera tarea será generar programáticamente las celdas para tu hoja de cálculo.

El objeto global `window` representa la ventana (o pestaña) del navegador. Tiene una propiedad `onload` que te permite definir el comportamiento cuando la ventana haya cargado toda la página, incluidas hojas de estilos y scripts.

Empieza configurando la propiedad `onload` de `window` a una función flecha sin parámetros. En la función, declara una variable `container` y asignale el valor de obtener el elemento por el `id` de `"container"`.

# --hints--

Debes acceder a la propiedad `onload` del objeto `window`.

```js
assert.match(code, /window\.onload/);
```

Debes configurar la propiedad `onload` a una función.

```js
assert.isFunction(window.onload);
```

Debes usar la sintaxis de flecha.

```js
assert.match(code, /window\.onload\s*=\s*\(/);
```

Tu función `onload` no debe tomar ningún parámetro.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)/);
```

Debes declarar una variable `container` en tu función `onload`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*(?:let|var|const)\s+container/);
```

Tu variable `container` debe ser declarada con `const`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container/);
```

Debes usar `document.getElementById()`

```js
assert.match(code, /document\.getElementById\(/);
```

Debes seleccionar el elemento con el `id` de `"container"`.

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

Debes asignar el elemento `#container` a `container`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
