---
id: 642e02be7845f13b014cd2b0
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Ahora que tienes una función `range`, también puedes usarla para crear un rango de letras.

Declara una función `charRange` usando `const` y sintaxis de flecha. Debe tomar un parámetro `start` y `end`. La función debe devolver implícitamente el resultado de llamar a `range()` con `start` y `end` como argumentos.

# --hints--

Deberías declarar una variable `charRange`.

```js
assert.match(code, /(?:let|const|var)\s+charRange/);
```

Tu variable `charRange` debe declararse con `const`.

```js
assert.match(code, /const\s+charRange/);
```

Tu variable `charRange` debe ser una función.

```js
assert.isFunction(charRange);
```

Tu función `charRange` debe usar sintaxis de flecha.

```js
assert.match(code, /const\s+charRange\s*=\s*\(.*\)\s*=>/);
```

Tu función `charRange` debe tomar `start` como primer parámetro.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start/);
```

Tu función `charRange` debe tomar `end` como segundo parámetro.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

Tu función `charRange` debe usar un retorno implícito.

```js
assert.notMatch(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

Tu función `charRange` debe llamar a tu función `range`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(/);
```

Debes pasar `start` y `end` como argumentos a tu llamada `range`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\s*,\s*end\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
