---
id: 6434750c53db16218f41e6e1
title: Paso 15
challengeType: 0
dashedName: step-15
---

# --description--

Recuerda que `range()` devuelve un array, por lo que puedes encadenar métodos de array directamente a la llamada de la función.

Llama a `range()` con `1` y `99` como argumentos, y encadena el método `.forEach()`. Pasa al método `.forEach()` un callback vacío que toma `number` como parámetro.

# --hints--

Deberías llamar a tu función `range()`.

```js
assert.lengthOf(code.match(/range\(/g), 2);
```

Deberías pasar `1` como primer argumento a tu llamada de `range()`.

```js
assert.match(code, /range\(\s*1/);
```

Deberías pasar `99` como segundo argumento a tu llamada de `range()`.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)/);
```

Deberías encadenar el método `.forEach()` a tu llamada de `range()`.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(/);
```

Deberías pasar una función callback a `.forEach()` usando la sintaxis de flecha.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

Tu función callback debería tener `number` como único parámetro.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\(\s*number\s*\)|number)\s*=>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
--fcc-editable-region--

--fcc-editable-region--
}
```
