---
id: 643475e13dc727231acd0f72
title: Paso 17
challengeType: 0
dashedName: step-17
---

# --description--

Ahora en tu llamada anidada a `.forEach()`, declara una variable `input`. Utiliza el método `.createElement()` del objeto `document` para crear un elemento `input`. Establece el atributo `type` a `"text"` y el atributo `id` a `letter + number`.

# --hints--

Debes declarar una variable `input`.

```js
assert.match(code, /(?:var|let|const)\s+input/)
```

Debes utilizar `const` para declarar tu variable `input`.

```js
assert.match(code, /const\s+input/)
```

Debes llamar al método `.createElement()` del objeto `document`.

```js
assert.lengthOf(code.match(/document\.createElement\(/g), 2)
```

Debes pasar la cadena `"input"` al método `.createElement()`.

```js
assert.match(code, /document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

Debes asignar tu nuevo elemento `input` a tu variable `input`.

```js
assert.match(code, /const\s+input\s*=\s*document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

Debes acceder a la propiedad `type` de tu elemento `input`.

```js
assert.match(code, /input\.type/);
```

Debes establecer el atributo `type` de tu elemento `input` a `"text"`.

```js
assert.match(code, /input\.type\s*=\s*('|"|`)text\1/)
```

Debes acceder a la propiedad `id` de tu elemento `input`.

```js
assert.match(code, /input\.id/);
```

Debes establecer el atributo `id` de tu elemento `input` a `letter + number`.

```js
assert.match(code, /input\.id\s*=\s*letter\s\+\snumber/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
--fcc-editable-region--
    letters.forEach(letter => {

    })
--fcc-editable-region--
  })
}
```
