---
id: 6437133052eaf04d7300e622
title: Paso 22
challengeType: 0
dashedName: step-22
---

# --description--

Declarar una función `average` que tome un arreglo de números como el parámetro `nums`. Debe devolver el promedio de todos los números en el arreglo.

El promedio se puede calcular dividiendo la suma de todos los números en el arreglo por la longitud del arreglo. Recuerda que tienes una función `sum` que puedes usar.

# --hints--

Debes declarar una variable `average`.

```js
assert.match(code, /(?:let|const|var)\s+average/);
```

Debes usar `const` para declarar la variable `average`.

```js
assert.match(code, /const\s+average/);
```

Tu variable `average` debe ser una función.

```js
assert.isFunction(average);
```

Tu función `average` debe usar la sintaxis de flecha.

```js
assert.match(code, /const\s+average\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

Tu función `average` debe tener un parámetro `nums`.

```js
assert.match(code, /const\s+average\s*=\s*(\(\s*nums\s*\)|nums)/);
```

Tu función `average` debe usar un retorno implícito.

```js
assert.notMatch(code, /const\s+average\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*{/);
```

Tu función `average` debe devolver el valor promedio del arreglo `nums`.

```js
assert.equal(average([1,2,3]), 2);
assert.equal(average([1,2,3,4,5]), 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
