---
id: 64496df724dd3716a71fe971
title: Paso 26
challengeType: 0
dashedName: step-26
---

# --description--

Las propiedades de los objetos consisten en pares clave/valor. Puedes usar nombres de propiedades abreviados al declarar un literal de objeto. Al utilizar la sintaxis de nombre de propiedad abreviado, el nombre de la variable se convierte en la clave de propiedad y su valor en el valor de propiedad.

El siguiente ejemplo declara un objeto `user` con las propiedades `userId`, `firstName` y `loggedIn`.

```js
const userId = 1;
const firstName = "John";
const loggedIn = true;

const user = {
  userId,
  firstName,
  loggedIn,
};

console.log(user); // { userId: 1, firstName: 'John', loggedIn: true }
```

Para hacer un seguimiento de todas las funciones de tu hoja de cálculo, declara un objeto `spreadsheetFunctions`. Usando la sintaxis de notación abreviada, establece `sum`, `average` y `median` como propiedades en el objeto `spreadsheetFunctions`.

# --hints--

Deberías declarar una variable `spreadsheetFunctions`.

```js
assert.match(code, /(?:const|let|var)\s+spreadsheetFunctions/);
```

Deberías usar `const` para declarar tu variable `spreadsheetFunctions`.

```js
assert.match(code, /const\s+spreadsheetFunctions/);
```

Tu variable `spreadsheetFunctions` debería ser un objeto.

```js
assert.isObject(spreadsheetFunctions);
```

Tu objeto `spreadsheetFunctions` debería tener una propiedad `sum`.

```js
assert.property(spreadsheetFunctions, "sum");
```

Tu propiedad `sum` debería ser tu función `sum`.

```js
assert.equal(spreadsheetFunctions?.sum, sum);
```

Tu objeto `spreadsheetFunctions` debería tener una propiedad `average`.

```js
assert.property(spreadsheetFunctions, "average");
```

Tu propiedad `average` debería ser tu función `average`.

```js
assert.equal(spreadsheetFunctions?.average, average);
```

Tu objeto `spreadsheetFunctions` debería tener una propiedad `median`.

```js
assert.property(spreadsheetFunctions, "median");
```

Tu propiedad `median` debería ser tu función `median`.

```js
assert.equal(spreadsheetFunctions?.median, median);
```

Deberías usar la sintaxis de desestructuración para asignar tus propiedades.

```js
const objectText = code.replace(/.*const\s+spreadsheetFunctions\s*=\s*\{([^}]*)}.*/s, "$1");
assert.include(objectText, "sum");
assert.include(objectText, "average");
assert.include(objectText, "median");
assert.notInclude(objectText, ":");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

const median = nums => {
  const sorted = nums.slice().sort((a, b) => a - b);
  const length = sorted.length;
  const middle = length / 2 - 1;
  return isEven(length)
    ? average([sorted[middle], sorted[middle + 1]])
    : sorted[Math.ceil(middle)];
}

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
