---
id: 66db529d37ad966480ebb633
title: 점심 메뉴 선택 프로그램 만들기
challengeType: 26
dashedName: build-a-lunch-picker-program
---

# --description--

이 실습에서는 점심 메뉴 관리를 돕는 프로그램을 만듭니다. 점심 메뉴 배열을 다루고, 배열에 항목을 추가하거나 제거하며, 무작위로 점심 메뉴를 선택하는 작업을 하게 됩니다.

**목표:** 아래 사용자 스토리를 모두 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 점심 메뉴 항목을 저장할 빈 배열을 할당한 `lunches` 변수를 만들어야 합니다.

2. 첫 번째 인수로 배열을, 두 번째 인수로 문자열을 받는 `addLunchToEnd` 함수를 만들어야 합니다. 이 함수는:
    - 문자열을 배열 끝에 추가해야 합니다.
    - 함수에 전달된 문자열 `[Lunch Item] added to the end of the lunch menu.`를 포함하는 문자열 `[Lunch Item]`를 콘솔에 출력해야 합니다.
    - 업데이트된 배열을 반환해야 합니다.

3. 첫 번째 인수로 배열을, 두 번째 인수로 문자열을 받는 `addLunchToStart` 함수를 만들어야 합니다. 이 함수는:
    - 문자열을 배열 시작 부분에 추가해야 합니다.
    - 함수에 전달된 문자열 `[Lunch Item] added to the start of the lunch menu.`를 포함하는 문자열 `[Lunch Item]`를 콘솔에 출력해야 합니다.
    - 업데이트된 배열을 반환해야 합니다.

4. 배열을 인수로 받는 `removeLastLunch` 함수를 만들어야 합니다. 이 함수는:
   - 배열의 마지막 요소를 제거해야 합니다.
   - 제거가 성공하면, 제거된 요소 `[Lunch Item] removed from the end of the lunch menu.`를 포함하는 문자열 `[Lunch Item]`를 콘솔에 출력해야 합니다.
   - 배열이 비어 있으면 문자열 `"No lunches to remove."`을 콘솔에 출력해야 합니다.
   - 업데이트된 배열을 반환해야 합니다.

5. 배열을 인수로 받는 `removeFirstLunch` 함수를 만들어야 합니다. 이 함수는:
   - 배열의 첫 번째 요소를 제거해야 합니다.
   - 제거가 성공하면, 제거된 요소 `[Lunch Item] removed from the start of the lunch menu.`를 포함하는 문자열 `[Lunch Item]`를 콘솔에 출력해야 합니다.
   - 배열이 비어 있으면 문자열 `"No lunches to remove."`를 콘솔에 출력해야 합니다.
   - 업데이트된 배열을 반환해야 합니다.

6. 배열을 인수로 받는 `getRandomLunch` 함수를 만들어야 합니다. 이 함수는:
   - 배열에서 무작위 요소를 선택해야 합니다.
   - 선택이 성공하면, 무작위 요소 `Randomly selected lunch: [Lunch Item]`를 포함하는 문자열 `[Lunch Item]`를 콘솔에 출력해야 합니다.
   - 배열이 비어 있으면 문자열 `"No lunches available."`를 콘솔에 출력해야 합니다.

7. 배열을 인수로 받는 `showLunchMenu` 함수를 만들어야 합니다. 이 함수는:
   - 배열에 요소가 있으면, 배열의 각 요소 `Menu items: [Lunch Item], [Lunch Item]...`를 순서대로 포함하는 문자열 `[Lunch item]`를 콘솔에 출력해야 합니다.
   - 배열이 비어 있으면 문자열 `"The menu is empty."`를 콘솔에 출력해야 합니다.

# --hints--

점심 메뉴 항목을 저장할 빈 배열을 할당한 `lunches` 변수를 선언해야 합니다.

```js
assert.isArray(lunches);
```

`addLunchToEnd` 함수를 정의해야 합니다.

```js
assert.isFunction(addLunchToEnd);
```

`addLunchToEnd` 함수는 두 개의 매개변수를 가져야 합니다.

```js
assert.lengthOf(addLunchToEnd, 2);
```

`addLunchToEnd(lunches, "Tacos")`는 문자열 `"Tacos added to the end of the lunch menu."`를 콘솔에 출력해야 합니다.

```js
const tempArr = [];
const temp = console.log;
const testLunches = []
try {
  console.log = obj => tempArr.push(obj);
  addLunchToEnd(testLunches, "Tacos");
  assert.strictEqual(tempArr[0], "Tacos added to the end of the lunch menu.");
  addLunchToEnd(testLunches, "Pizza");
  assert.strictEqual(tempArr[1], "Pizza added to the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToEnd(["Pizza", "Tacos"], "Burger")`는 `["Pizza", "Tacos", "Burger"]`를 반환해야 합니다.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToEnd(["Pizza", "Tacos"], "Burger"), ["Pizza", "Tacos", "Burger"]);
  assert.deepStrictEqual(addLunchToEnd(["Fries"], "Tacos"), ["Fries", "Tacos"]);
} finally {
  console.log = temp;
}
```

`addLunchToStart` 함수를 정의해야 합니다.

```js
assert.isFunction(addLunchToStart);
```

`addLunchToStart` 함수는 두 개의 매개변수를 가져야 합니다.

```js
assert.lengthOf(addLunchToStart, 2);
```

`addLunchToStart(lunches, "Sushi")`는 문자열 `"Sushi added to the start of the lunch menu."`를 콘솔에 출력해야 합니다.

```js
const tempArr = [];
const testLunches = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  addLunchToStart(testLunches, "Sushi");
  assert.strictEqual(tempArr[0], "Sushi added to the start of the lunch menu.");  
  addLunchToStart(testLunches, "Burger");
  assert.strictEqual(tempArr[1], "Burger added to the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToStart(["Burger", "Sushi"], "Pizza")`는 `["Pizza", "Burger", "Sushi"]`를 반환해야 합니다.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToStart(["Burger", "Sushi"], "Pizza"), ["Pizza", "Burger", "Sushi"]);
  assert.deepStrictEqual(addLunchToStart(["Noodles"], "Sushi"), ["Sushi", "Noodles"]);
} finally {
  console.log = temp;
}
```

`removeLastLunch` 함수를 정의해야 합니다.

```js
assert.isFunction(removeLastLunch)
```

`removeLastLunch` 함수는 하나의 매개변수를 가져야 합니다.

```js
assert.lengthOf(removeLastLunch, 1);
```

입력 배열이 비어 있을 때, `removeLastLunch` 함수는 문자열 `"No lunches to remove."`를 콘솔에 출력해야 합니다.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Stew", "Soup", "Toast"])`는 문자열 `"Toast removed from the end of the lunch menu."`를 콘솔에 출력해야 합니다.

```js
let testLunches = ["Stew", "Soup", "Toast"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "Toast removed from the end of the lunch menu.");
  testLunches = ["Rice", "Pizza"];
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[1], "Pizza removed from the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Sushi", "Pizza", "Noodles"])`는 `["Sushi", "Pizza"]`를 반환해야 합니다.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeLastLunch(["Sushi", "Pizza", "Noodles"]), ["Sushi", "Pizza"]);
  assert.deepStrictEqual(removeLastLunch(["Pizza", "Burger"]), ["Pizza"]);
} finally {
  console.log = temp;
}
```

`removeFirstLunch` 함수를 정의해야 합니다.

```js
assert.isFunction(removeFirstLunch);
```

`removeFirstLunch` 함수는 하나의 매개변수를 가져야 합니다.

```js
assert.lengthOf(removeFirstLunch, 1);
```

입력 배열이 비어 있을 때, `removeFirstLunch` 함수는 문자열 `"No lunches to remove."`를 콘솔에 출력해야 합니다.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Salad", "Eggs", "Cheese"])`는 문자열 `"Salad removed from the start of the lunch menu."`를 콘솔에 출력해야 합니다.

```js
let testLunches = ["Salad", "Eggs", "Cheese"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "Salad removed from the start of the lunch menu.");
  testLunches = ["Stew", "Soup", "Toast"]
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[1], "Stew removed from the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Sushi", "Pizza", "Burger"])`는 `["Pizza", "Burger"]`를 반환해야 합니다.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeFirstLunch(["Sushi", "Pizza", "Burger"]), ["Pizza", "Burger"]);
  assert.deepStrictEqual(removeFirstLunch(["Pizza", "Burger"]), ["Burger"]);
} finally {
  console.log = temp;
}
```

`addLunchToEnd`, `addLunchToStart`, `removeLastLunch`, 그리고 `removeFirstLunch`는 업데이트 후 인수로 받은 동일한 배열을 반환해야 합니다.

```js
const temp = console.log;
console.log = () => {}
const testLunches = ["Spam"];
try {
  assert.equal(addLunchToEnd(testLunches, "Pizza"), testLunches);
  assert.equal(addLunchToStart(testLunches, "Caviar"), testLunches);
  assert.equal(removeLastLunch(testLunches), testLunches);
  assert.equal(removeFirstLunch(testLunches), testLunches);
} finally {
  console.log = temp;
}
```

`getRandomLunch` 함수를 정의해야 합니다.

```js
assert.isFunction(getRandomLunch);
```

`getRandomLunch` 함수는 하나의 매개변수를 가져야 합니다.

```js
assert.lengthOf(getRandomLunch, 1);
```

입력 배열이 비어 있을 때, `getRandomLunch` 함수는 문자열 `"No lunches available."`를 콘솔에 출력해야 합니다.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  getRandomLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches available.");
} finally {
  console.log = temp;
}
```

입력 배열이 비어 있지 않을 때, `getRandomLunch` 함수는 `Randomly selected lunch: [Lunch Item]` 형식의 문자열을 콘솔에 출력해야 합니다.

```js
const testLunches = ["Fish", "Fries", "Roast"];
const tempRandom = Math.random;
const tempArr = [];
const temp = console.log;

try {
  console.log = obj => tempArr.push(obj);
  // check that it correctly outputs the first item
  Math.random = () => 0;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[0], `Randomly selected lunch: ${testLunches[0]}`);

  // Second item
  Math.random = () => 0.4;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[1], `Randomly selected lunch: ${testLunches[1]}`);

  // Third item
  Math.random = () => 0.8;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[2], `Randomly selected lunch: ${testLunches[2]}`);

} finally {
  Math.random = tempRandom;
  console.log = temp;
}
```

`getRandomLunch` 함수는 인수로 받은 배열을 변경하지 않아야 합니다.

```js
const testLunches = ["Fish", "Fries", "Roast"];
let originalLunches = [...testLunches];
const tempRandom = Math.random;

try {
  Math.random = () => 0;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.4;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.8;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  const emptyLunches = [];
  let originalEmpty = [...emptyLunches];
  getRandomLunch(emptyLunches);
  assert.sameOrderedMembers(emptyLunches, originalEmpty);

} finally {
  Math.random = tempRandom;
}
```

`showLunchMenu` 함수를 정의해야 합니다.

```js
assert.isFunction(showLunchMenu);
```

`showLunchMenu` 함수는 하나의 매개변수를 가져야 합니다.

```js
assert.lengthOf(showLunchMenu, 1);
```

입력 배열이 비어 있을 때, `showLunchMenu` 함수는 문자열 `"The menu is empty."`를 콘솔에 출력해야 합니다.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], "The menu is empty.");
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Greens", "Corns", "Beans"])`는 `"Menu items: Greens, Corns, Beans"`를 콘솔에 출력해야 합니다.

```js
const testLunches = ["Greens", "Corns", "Beans"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Pizza", "Burger", "Fries", "Salad"])`는 `"Menu items: Pizza, Burger, Fries, Salad"`를 콘솔에 출력해야 합니다.

```js
const testLunches = ["Pizza", "Burger", "Fries", "Salad"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const lunches = [];

// Add a new lunch to the end of the list
function addLunchToEnd(lunches, newLunch) {
  lunches.push(newLunch);
  console.log(`${newLunch} added to the end of the lunch menu.`);
  return lunches;
}

// Add a new lunch to the start of the list
function addLunchToStart(lunches, newLunch) {
  lunches.unshift(newLunch);
  console.log(`${newLunch} added to the start of the lunch menu.`);
  return lunches;
}

// Remove the last lunch from the list
function removeLastLunch(lunches) {
  const removedLunch = lunches.pop();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the end of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Remove the first lunch from the list
function removeFirstLunch(lunches) {
  const removedLunch = lunches.shift();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the start of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Function to get a random lunch from the list
function getRandomLunch(lunches) {
  if (lunches.length === 0) {
    console.log("No lunches available.");
  }
  else {
    const randomIndex = Math.floor(Math.random() * lunches.length);
    const randomLunch = lunches[randomIndex];
    console.log(`Randomly selected lunch: ${randomLunch}`);
  }
}

// Function to display all the lunches in the list
function showLunchMenu(lunches) {
  if (lunches.length === 0) {
    console.log("The menu is empty.");
  }
  else {
    console.log(`Menu items: ${lunches.join(', ')}`);
  }
}

addLunchToEnd(lunches, "Tacos");
addLunchToStart(lunches, "Sushi");
removeLastLunch(lunches);
removeFirstLunch(lunches);
getRandomLunch(lunches);
showLunchMenu(lunches);

```
