---
id: 587d8248367417b2b2512c3b
title: منع IE من فتح HTML غير موثوق به باستخدام helmet.ieNoOpen()
challengeType: 2
forumTopicId: 301584
dashedName: prevent-ie-from-opening-untrusted-html-with-helmet-ienoopen
---

# --description--

كتذكير، يُبنى هذا المشروع على قالب جاهز تم نسخه من <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

بعض تطبيقات الويب تقدم HTML غير موثوق به للتحميل. بعض إصدارات Internet Explorer تفتح هذه الملفات HTML افتراضيًا في سياق موقعك. هذا يعني أن صفحة HTML غير موثوق بها قد تبدأ بتنفيذ أفعال ضارة في سياق صفحاتك. تضبط برمجيات الوسيطة هذه القسم العلوي X-Download-Options على noopen. هذا يمنع مستخدمي IE من تنفيذ التنزيلات في سياق الموقع الموثوق.

# --instructions--

استخدم الطريقة `helmet.ieNoOpen()` على الخادم الخاص بك.

# --hints--

يجب تركيب برمجيات الوسيطة helmet.ieNoOpen() بشكل صحيح

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'ienoopen');
assert.equal(data.headers['x-download-options'], 'noopen');
```

