---
id: 587d8247367417b2b2512c38
title: تخفيف خطر النقر الاحتيالي باستخدام helmet.frameguard()
challengeType: 2
forumTopicId: 301582
dashedName: mitigate-the-risk-of-clickjacking-with-helmet-frameguard
---

# --description--

كتذكير، يُبنى هذا المشروع على قالب جاهز تم نسخه من <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

قد تُوضع صفحتك داخل `<frame>` أو `<iframe>` بدون موافقتك. هذا قد يؤدي إلى هجمات النقر الاحتيالي، من أمور أخرى. النقر الاحتيالي هو تقنية لخداع المستخدم للتفاعل مع صفحة مختلفة عما يعتقده. يحدث ذلك بتنفيذ صفحتك في سياق خبيث عبر الإطار. في هذا السياق، يمكن للقراصنة وضع طبقة مخفية فوق صفحتك. تُستخدم أزرار مخفية لتشغيل شفرات ضارة. هذه البرمجيات الوسيطة تضبط القسم العلوي X-Frame-Options. تحدد من يمكنه وضع موقعك في إطار. لها ثلاثة أوضاع: DENY، SAMEORIGIN، وALLOW-FROM.

لا نحتاج إلى أن تُعرض تطبيقنا داخل إطار.

# --instructions--

استخدم `helmet.frameguard()` مع تمرير كائن التهيئة `{action: 'deny'}`.

# --hints--

يجب تركيب برمجيات وسيطة helmet.frameguard() بشكل صحيح

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(
  data.appStack,
  'frameguard',
  'helmet.frameguard() middleware is not mounted correctly'
);
```

يجب ضبط 'action' في helmet.frameguard() على 'DENY'

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.property(data.headers, 'x-frame-options');
assert.equal(data.headers['x-frame-options'], 'DENY');
```

