---
id: 587d8248367417b2b2512c3d
title: تعطيل جلب سجلات DNS المسبق باستخدام helmet.dnsPrefetchControl()
challengeType: 2
forumTopicId: 301577
dashedName: disable-dns-prefetching-with-helmet-dnsprefetchcontrol
---

# --description--

كتذكير، يُبنى هذا المشروع على قالب جاهز تم نسخه من <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

لتحسين الأداء، تقوم معظم المتصفحات بجلب سجلات DNS مسبقًا للروابط في الصفحة. بهذه الطريقة يكون عنوان الـ IP الوجهة معروفًا عند نقر المستخدم على رابط. قد يؤدي هذا إلى استخدام مفرط لخدمة DNS (إذا كنت تملك موقعًا كبيرًا يزوره ملايين الأشخاص...)، أو مشكلات خصوصية (قد يستنتج متنصت أنك على صفحة معينة)، أو تغيير إحصائيات الصفحة (قد تظهر بعض الروابط كأنها زُرت حتى لو لم تُزر). إذا كانت لديك احتياجات أمان عالية، يمكنك تعطيل جلب سجلات DNS المسبق، مع تحمل خسارة في الأداء.

# --instructions--

استخدم الطريقة `helmet.dnsPrefetchControl()` على الخادم الخاص بك.

# --hints--

يجب تركيب برمجيات وسيطة helmet.dnsPrefetchControl() بشكل صحيح

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'dnsPrefetchControl');
assert.equal(data.headers['x-dns-prefetch-control'], 'off');
```

