---
id: 587d8247367417b2b2512c36
title: تثبيت واستدعاء Helmet
challengeType: 2
forumTopicId: 301581
dashedName: install-and-require-helmet
---

# --description--

سيتطلب منك العمل على هذه التحديات كتابة شفرتك باستخدام إحدى الطرق التالية:

- استنسخ <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">مستودع GitHub هذا</a> وأكمل هذه التحديات محليًا.
- استخدم منشئ مواقع من اختيارك لإكمال المشروع. تأكد من دمج جميع الملفات من مستودع GitHub الخاص بنا.

يساعدك Helmet على تأمين تطبيقات Express الخاصة بك عن طريق تعيين عدة رؤوس HTTP.

# --instructions--

تذهب كل شفرتك لهذه الدروس في ملف `myApp.js` بين أسطر الشفرة التي بدأناها لك. لا تغير أو تحذف الشفرة التي أضفناها لك.

تم تثبيت إصدار Helmet `3.21.3` بالفعل، لذا استدعِه كـ `helmet` في `myApp.js`.

# --hints--

يجب أن يكون إصدار `helmet` `3.21.3` موجودًا في `package.json`

```js
const response = await fetch(code + '/_api/package.json');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
const packJson = JSON.parse(data);
const helmet = packJson.dependencies.helmet;
assert(helmet === '3.21.3' || helmet === '^3.21.3');
```

