---
id: 587d8249367417b2b2512c3e
title: تعطيل التخزين المؤقت على جانب العميل باستخدام helmet.noCache()
challengeType: 2
forumTopicId: 301576
dashedName: disable-client-side-caching-with-helmet-nocache
---

# --description--

كتذكير، يُبنى هذا المشروع على قالب جاهز تم نسخه من <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

إذا كنت تصدر تحديثًا لموقعك، وترغب في أن يقوم المستخدمون دائمًا بتنزيل الإصدار الأحدث، يمكنك محاولة تعطيل التخزين المؤقت في متصفح العميل. هذا مفيد أيضًا في التطوير. التخزين المؤقت يحسن الأداء، والذي ستفقده، لذا استخدم هذا الخيار فقط عند وجود حاجة حقيقية.

# --instructions--

استخدم الطريقة `helmet.noCache()` على الخادم الخاص بك.

# --hints--

يجب تركيب برمجيات وسيطة helmet.noCache() بشكل صحيح

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nocache');
assert.equal(
  data.headers['cache-control'],
  'no-store, no-cache, must-revalidate, proxy-revalidate'
);
```

