---
id: 587d8248367417b2b2512c3a
title: تجنب استنتاج نوع MIME للاستجابة باستخدام helmet.noSniff()
challengeType: 2
forumTopicId: 301574
dashedName: avoid-inferring-the-response-mime-type-with-helmet-nosniff
---

# --description--

كتذكير، يُبنى هذا المشروع على قالب جاهز تم نسخه من <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

يمكن للمتصفحات استخدام تحليل المحتوى أو نوع MIME لتجاوز القسم العلوي `Content-Type` للاستجابة بهدف تخمين ومعالجة البيانات باستخدام نوع محتوى ضمني. رغم أن هذا قد يكون مريحًا في بعض الحالات، إلا أنه قد يؤدي إلى هجمات خطيرة. تقوم برمجيات وسيطة هذه بتعيين القسم العلوي `X-Content-Type-Options` إلى `nosniff`، مما يوجه المتصفح بعدم تجاوز `Content-Type` المقدم.

# --instructions--

استخدم الطريقة `helmet.noSniff()` على الخادم الخاص بك.

# --hints--

يجب تركيب برمجيات وسيطة helmet.noSniff() بشكل صحيح

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nosniff');
assert.equal(data.headers['x-content-type-options'], 'nosniff');
```

