---
id: 587d8247367417b2b2512c36
title: تثبيت واستدعاء Helmet
challengeType: 2
forumTopicId: 301581
dashedName: install-and-require-helmet
---

# --description--

العمل على هذه التحديات سيتطلب منك كتابة كودك باستخدام إحدى الطرق التالية:

- استنسخ <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">مستودع GitHub هذا</a> وأكمل هذه التحديات محليًا.
- استخدم منشئ مواقع من اختيارك لإكمال المشروع. تأكد من دمج جميع الملفات من مستودع GitHub الخاص بنا.

يساعدك Helmet على تأمين تطبيقات Express الخاصة بك عن طريق تعيين عدة رؤوس HTTP.

# --instructions--

كل الكود الخاص بهذه الدروس يذهب في ملف `myApp.js` بين أسطر الكود التي بدأناها لك. لا تغير أو تحذف الكود الذي أضفناه لك.

تم تثبيت إصدار Helmet `3.21.3` بالفعل، لذا استدعِه كـ `helmet` في `myApp.js`.

# --hints--

يجب أن يكون إصدار `helmet` `3.21.3` موجودًا في `package.json`

```js
const response = await fetch(code + '/_api/package.json');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
const packJson = JSON.parse(data);
const helmet = packJson.dependencies.helmet;
assert(helmet === '3.21.3' || helmet === '^3.21.3');
```

