---
id: 587d8249367417b2b2512c40
title: "تهيئة Helmet باستخدام البرمجيات الوسيطة ‘parent’ helmet()"
challengeType: 2
forumTopicId: 301575
dashedName: configure-helmet-using-the-parent-helmet-middleware
---

# --description--

كتذكير، يُبنى هذا المشروع على مشروع بداية تم استنساخه من <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

سيشمل `app.use(helmet())` تلقائيًا كل برمجيات وسيطة تم تقديمها أعلاه، باستثناء `noCache()` و`contentSecurityPolicy()`، لكن يمكن تفعيلها إذا لزم الأمر. يمكنك أيضًا تعطيل أو تكوين أي برمجيات وسيطة أخرى بشكل فردي باستخدام كائن التكوين.

**مثال:**

```js
app.use(helmet({
  frameguard: {         // configure
    action: 'deny'
  },
  contentSecurityPolicy: {    // enable and configure
    directives: {
      defaultSrc: ["'self'"],
      styleSrc: ['style.com'],
    }
  },
  dnsPrefetchControl: false     // disable
}))
```

قدمنا كل برمجيات وسيطة بشكل منفصل لأغراض التعليم ولتسهيل الاختبار. استخدام برمجيات وسيطة `helmet()` الأصلية سهل التنفيذ في مشروع حقيقي.

# --hints--

لا توجد اختبارات - هذا تحدٍّ وصفي

```js
assert(true);
```

