---
id: 587d8247367417b2b2512c38
title: تخفيف خطر النقر الاحتيالي باستخدام helmet.frameguard()
challengeType: 2
forumTopicId: 301582
dashedName: mitigate-the-risk-of-clickjacking-with-helmet-frameguard
---

# --description--

كتذكير، يُبنى هذا المشروع على مشروع بداية تم استنساخه من <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

قد تُوضع صفحتك في `<frame>` أو `<iframe>` بدون موافقتك. هذا قد يؤدي إلى هجمات clickjacking، من بين أمور أخرى. clickjacking هي تقنية لخداع المستخدم للتفاعل مع صفحة مختلفة عما يعتقده المستخدم. يحدث ذلك بتنفيذ صفحتك في سياق خبيث، عبر تضمينها في إطار. في هذا السياق، يمكن للمخترق وضع طبقة مخفية فوق صفحتك. يمكن استخدام أزرار مخفية لتشغيل سكربتات ضارة. هذه البرمجيات الوسيطة تضبط رأس X-Frame-Options. تحدد من يمكنه وضع موقعك في إطار. لها ثلاثة أوضاع: DENY، SAMEORIGIN، وALLOW-FROM.

لا نحتاج إلى أن تُعرض تطبيقنا داخل إطار.

# --instructions--

استخدم `helmet.frameguard()` مع تمرير كائن التهيئة `{action: 'deny'}`.

# --hints--

يجب تركيب برمجيات وسيطة helmet.frameguard() بشكل صحيح

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(
  data.appStack,
  'frameguard',
  'helmet.frameguard() middleware is not mounted correctly'
);
```

يجب ضبط 'action' في helmet.frameguard() على 'DENY'

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.property(data.headers, 'x-frame-options');
assert.equal(data.headers['x-frame-options'], 'DENY');
```

