---
id: 587d8247367417b2b2512c37
title: إخفاء المعلومات التي قد تكون خطيرة باستخدام helmet.hidePoweredBy()
challengeType: 2
forumTopicId: 301580
dashedName: hide-potentially-dangerous-information-using-helmet-hidepoweredby
---

# --description--

كتذكير، يُبنى هذا المشروع على قالب جاهز تم نسخه من <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

يمكن للقراصنة استغلال الثغرات المعروفة في Express/Node إذا رأوا أن موقعك يعمل بواسطة Express. يُرسل `X-Powered-By: Express` في كل طلب صادر من Express افتراضيًا. استخدم برمجيات وسيطة `helmet.hidePoweredBy()` لإزالة القسم العلوي X-Powered-By.

# --hints--

يجب تركيب برمجيات وسيطة helmet.hidePoweredBy() بشكل صحيح

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hidePoweredBy');
assert.notEqual(data.headers['x-powered-by'], 'Express');
```

