---
id: 587d8247367417b2b2512c39
title: تقليل خطر هجمات البرمجة النصية عبر المواقع (XSS) باستخدام helmet.xssFilter()
  Mitigate the Risk of Cross Site Scripting (XSS) Attacks with
  helmet.xssFilter()
challengeType: 2
forumTopicId: 301583
dashedName: mitigate-the-risk-of-cross-site-scripting-xss-attacks-with-helmet-xssfilter
---

# --description--

كتذكير، يُبنى هذا المشروع على قالب جاهز تم نسخه من <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

البرمجة النصية عبر المواقع (XSS) هي نوع شائع من الهجمات حيث تُحقن نصوص خبيثة في صفحات ضعيفة بهدف سرقة بيانات حساسة مثل ملفات تعريف الارتباط للجلسة أو كلمات المرور.

القاعدة الأساسية لتقليل خطر هجوم XSS بسيطة: "لا تثق أبدًا بمُدخَل المستخدم". كمطوِّر يجب عليك دائمًا تنظيف كل المُدخَلات القادمة من الخارج. يشمل ذلك البيانات القادمة من النماذج، واستعلامات GET في عناوين URL، وحتى من محتوى POST. التنظيف يعني أنه يجب عليك إيجاد وترميز الأحرف التي قد تكون خطرة مثل &lt; و >.

يمكن للمتصفحات الحديثة المساعدة في تقليل الخطر من خلال اعتماد استراتيجيات برمجية أفضل. غالبًا ما يمكن تكوين هذه الاستراتيجيات عبر رؤوس HTTP.

رأس HTTP المسمى X-XSS-Protection هو حماية أساسية. يكتشف المتصفح نصًا محقونًا محتملًا باستخدام فلتر استدلالي. إذا كان الرأس مفعّلًا، يغير المتصفح شفرة النص ليعطلها. لا يزال هذا الدعم محدودًا.

# --instructions--

استخدم `helmet.xssFilter()` لتنظيف المُدخَل المرسل إلى خادمك.

# --hints--

يجب تركيب برمجيات وسيطة helmet.xssFilter() بشكل صحيح

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'xXssProtection');
assert.property(data.headers, 'x-xss-protection');
```

