---
id: 587d8248367417b2b2512c3c
title: اطلب من المتصفحات الوصول إلى موقعك عبر HTTPS فقط باستخدام helmet.hsts()
challengeType: 2
forumTopicId: 301573
dashedName: ask-browsers-to-access-your-site-via-https-only-with-helmet-hsts
---

# --description--

كتذكير، يُبنى هذا المشروع على مشروع بداية تم نسخه من <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

سياسة أمان النقل الصارم عبر HTTP (HSTS) هي سياسة أمان للويب تساعد في حماية المواقع من هجمات تخفيض البروتوكول واختطاف ملفات تعريف الارتباط. إذا كان يمكن الوصول إلى موقعك عبر HTTPS، يمكنك طلب من متصفحات المستخدمين تجنب استخدام HTTP غير الآمن. من خلال تعيين القسم العلوي Strict-Transport-Security، تخبر المتصفحات باستخدام HTTPS للطلبات المستقبلية خلال فترة زمنية محددة. هذا يعمل على الطلبات التي تأتي بعد الطلب الأولي.

# --instructions--

كوّن `helmet.hsts()` لاستخدام HTTPS للأيام الـ 90 القادمة. مرر كائن التكوين `{maxAge: timeInSeconds, force: true}`. يمكنك إنشاء متغير `ninetyDaysInSeconds = 90*24*60*60;` لاستخدامه في `timeInSeconds`.

ملاحظة: تكوين HTTPS على موقع مخصص يتطلب الحصول على نطاق وشهادة SSL/TLS.

# --hints--

يجب تركيب برمجيات وسيطة helmet.hsts() بشكل صحيح

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hsts');
assert.property(data.headers, 'strict-transport-security');
```

يجب أن تكون قيمة maxAge مساوية لـ 7776000 ثانية (90 يومًا)

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.match(
  data.headers['strict-transport-security'],
  /^max-age=7776000;?/
);
```

