---
id: 58a25bcff9fc0f352b528e7d
title: تجزئة ومقارنة كلمات المرور بشكل غير متزامن
challengeType: 2
forumTopicId: 301578
dashedName: hash-and-compare-passwords-asynchronously
---

# --description--

كتذكير، يُبنى هذا المشروع على مشروع بداية تم استنساخه من <a href="https://github.com/freeCodeCamp/boilerplate-bcrypt/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

بما أن التجزئة مصممة لتكون مكثفة حسابيًا، يُنصح بتنفيذها بشكل غير متزامن على الخادم لتجنب حجب الاتصالات الواردة أثناء التجزئة. كل ما عليك فعله لتجزئة كلمة مرور بشكل غير متزامن هو استدعاء

```js
bcrypt.hash(myPlaintextPassword, saltRounds, (err, hash) => {
  /*Store hash in your db*/
});
```

# --instructions--

أضف دالة التجزئة هذه إلى الخادم (لقد عرفنا المتغيرات المستخدمة في الدالة لتستخدمها) وسجلها في وحدة التحكم لتراها! في هذه المرحلة عادةً ما تحفظ التجزئة في قاعدة البيانات.

الآن عندما تحتاج إلى معرفة ما إذا كان المُدخَل الجديد هو نفس البيانات التي في التجزئة، تستخدم دالة المقارنة فقط.

```js
bcrypt.compare(myPlaintextPassword, hash, (err, res) => {
  /*res == true or false*/
});
```

أضف هذا إلى دالة التجزئة الموجودة لديك (لأنك تحتاج إلى انتظار اكتمال التجزئة قبل استدعاء دالة المقارنة) بعد تسجيل التجزئة المكتملة وسجل 'res' في وحدة التحكم داخل دالة المقارنة. سترى في وحدة التحكم التجزئة، ثم تُطبع 'true'! إذا غيرت 'myPlaintextPassword' في دالة المقارنة إلى 'someOtherPlaintextPassword'، فسيظهر false.

```js
bcrypt.hash('passw0rd!', 13, (err, hash) => {
  console.log(hash);
  //$2a$12$Y.PHPE15wR25qrrtgGkiYe2sXo98cjuMCG1YwSI5rJW1DSJp0gEYS
  bcrypt.compare('passw0rd!', hash, (err, res) => {
    console.log(res); //true
  });
});

```

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح.

# --hints--

يجب توليد التجزئة بشكل غير متزامن ومقارنتها بشكل صحيح.

```js
const response = await fetch(code + '/_api/server.js');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
assert.match(
  data,
  /START_ASYNC[^]*bcrypt.hash.*myPlaintextPassword( |),( |)saltRounds( |),( |).*err( |),( |)hash[^]*END_ASYNC/gi,
  'You should call bcrypt.hash on myPlaintextPassword and saltRounds and handle err and hash as a result in the callback'
);
assert.match(
  data,
  /START_ASYNC[^]*bcrypt.hash[^]*bcrypt.compare.*myPlaintextPassword( |),( |)hash( |),( |).*err( |),( |)res[^]*}[^]*}[^]*END_ASYNC/gi,
  'Nested within the hash function should be the compare function comparing myPlaintextPassword to hash'
);
```

