---
id: 58a25bcff9fc0f352b528e7e
title: تجزئة ومقارنة كلمات المرور بشكل متزامن
challengeType: 2
forumTopicId: 301579
dashedName: hash-and-compare-passwords-synchronously
---

# --description--

كتذكير، يُبنى هذا المشروع على مشروع بداية تم نسخه من <a href="https://github.com/freeCodeCamp/boilerplate-bcrypt/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

تجزئة البيانات بشكل متزامن سهلة التنفيذ أيضًا لكنها قد تسبب تباطؤًا إذا استُخدمت على الخادم مع تكلفة عالية أو مع تجزئة متكررة جدًا. التجزئة بهذه الطريقة سهلة مثل استدعاء

```js
var hash = bcrypt.hashSync(myPlaintextPassword, saltRounds);
```

أضف هذه الطريقة للتجزئة إلى شفرتك ثم سجّل النتيجة في وحدة التحكم. المتغيرات المستخدمة معرفة مسبقًا في الخادم، لذا لا تحتاج لتعديلها. قد تلاحظ أنه رغم تجزئتك لنفس كلمة المرور كما في الدالة غير المتزامنة، النتيجة في وحدة التحكم مختلفة — يعود ذلك إلى أن الملح يُولد عشوائيًا في كل مرة كما يظهر في أول 22 حرفًا من السلسلة النصية الثالثة في التجزئة. لمقارنة مُدخَل كلمة المرور مع التجزئة المتزامنة الجديدة، استخدم طريقة compareSync:

```js
var result = bcrypt.compareSync(myPlaintextPassword, hash);
```

والنتيجة تكون قيمة منطقية true أو false.

# --instructions--

أضف الدالة وسجّل النتيجة في وحدة التحكم لترى عملها.

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح.

# --hints--

يجب توليد التجزئة المتزامنة ومقارنتها بشكل صحيح.

```js
const response = await fetch(code + '/_api/server.js');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
assert.match(
  data,
  /START_SYNC[^]*hash.*=.*bcrypt.hashSync.*myPlaintextPassword( |),( |)saltRounds[^]*END_SYNC/gi,
  'You should call bcrypt.hashSync on myPlaintextPassword with saltRounds'
);
assert.match(
  data,
  /START_SYNC[^]*result.*=.*bcrypt.compareSync.*myPlaintextPassword( |),( |)hash[^]*END_SYNC/gi,
  'You should call bcrypt.compareSync on myPlaintextPassword with the hash generated in the last line'
);
```

