---
id: 7a97b79c096ecb000d03bd36
title: N번째 피보나치 수 계산기 만들기
challengeType: 26
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**목표:** 아래 사용자 스토리를 완성하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `fibonacci`라는 이름의 함수를 만들어야 합니다.
2. `sequence` 함수 내에 `fibonacci`라는 배열을 정의하고, `[0, 1]` 값으로 초기화해야 합니다.
3. `fibonacci` 함수는 하나의 매개변수, 양의 정수 `n`를 받아야 합니다.
4. `fibonacci(n)`를 호출하면 동적 프로그래밍 방식을 사용해 피보나치 수열에서 각 숫자가 바로 앞 두 숫자의 합인 `n`번째 수를 계산해 반환해야 합니다.
5. 피보나치 수열에서 `n` 위치에 있는 각 계산된 숫자는 `sequence` 배열의 `n - 1` 인덱스에 저장되어야 합니다.

# --hints--

`fibonacci`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(fibonacci);
```

`fibonacci` 함수는 하나의 매개변수를 받아야 합니다.

```js
const params = __helpers.getFunctionParams(fibonacci.toString());
assert.equal(params.length, 1);
```

`sequence` 함수 내에 `fibonacci`라는 배열이 있어야 하며, `[0, 1]`로 초기화되어야 합니다.

```js
assert.match(__helpers.removeJSComments(fibonacci.toString()), /(const|let|var)\s+sequence\s*=\s*\[\s*0\s*,\s*1\s*\]/);
```

`fibonacci(0)`는 `0`를 반환해야 합니다.

```js
assert.equal(fibonacci(0), 0);
```

`fibonacci(1)`는 `1`를 반환해야 합니다.

```js
assert.equal(fibonacci(1), 1);
```

`fibonacci(2)`는 `1`를 반환해야 합니다.

```js
assert.equal(fibonacci(2), 1);
```

`fibonacci(3)`는 `2`를 반환해야 합니다.

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)`는 `5`를 반환해야 합니다.

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)`는 `55`를 반환해야 합니다.

```js
assert.equal(fibonacci(10), 55);
```

`fibonacci(15)`는 `610`를 반환해야 합니다.

```js
assert.equal(fibonacci(15), 610);
```

코드에서 재귀를 사용하지 않아야 합니다.

```js
const bodyMatch = __helpers.removeJSComments(code).match(/function\s+fibonacci\s*\([^)]*\)\s*\{([\s\S]*)\}/);
assert(bodyMatch && !bodyMatch[1].match(/\bfibonacci\s*\(/));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fibonacci(n) {
  const sequence = [0, 1];

  if (n < 2) {
    return sequence[n];
  }

  for (let i = 2; i <= n; i++) {
    sequence.push(sequence[i - 1] + sequence[i - 2]);
  }

  return sequence[n];
}
```
