---
id: 60eebd07ea685b0e777b5583
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Welcome to the Registration Form project! Start by adding the `!DOCTYPE html` declaration at the top of the document so the browser knows what type of document it's reading.

# --hints--

Seu código deve conter a referência `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Você deve incluir um espaço após a referência `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Você deve definir o tipo de documento como `html`.

```js
assert(code.match(/html/gi));
```

Você deve fechar a declaração `DOCTYPE` com um `>` após o tipo.

```js
assert(code.match(/html\s*>/gi));
```

A declaração `DOCTYPE` deve estar no início do HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
