---
id: 62cc5b1779e4d313466f73c5
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Adicione um elemento `link` dentro do elemento `head`. Dê a ele um valor do atributo `rel` de `stylesheet` e um valor do atributo `href` de `styles.css`.

# --hints--

O código deve ter um elemento `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

Você não deve mudar as tags `head` existentes. Verifique se você não excluiu a tag de fechamento.

```js
const heads = document.querySelectorAll('head');
assert.strictEqual(heads?.length, 1);
```

Você deve ter um elemento `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

O elemento `link` deve ser um elemento nulo, sem uma tag de fechamento `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

O elemento `link` deve estar dentro do elemento `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
const link = document.querySelector('link');
assert.strictEqual(link?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
