---
id: 66be24cb4144f955b6bcc550
title: Costruisci una pagina di carte da gioco
challengeType: 25
dashedName: build-a-page-of-playing-cards
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Devi costruire una pagina web che mostri almeno tre carte da gioco.
1. Devi avere un elemento `main` con un ID `playing-cards`.
1. All'interno del tuo elemento `#playing-cards` devi avere almeno tre elementi `div`, ciascuno con una classe `card`.
1. All'interno di ogni elemento `.card` devi avere tre elementi `div`, il primo con una classe `left`, il secondo con una classe `middle` e il terzo con una classe `right`.
1. Il tuo elemento `#playing-cards` deve usare flexbox per centrare orizzontalmente i suoi figli, permettere loro di andare a capo e mettere uno spazio `20px` tra di loro.
1. Ciascuno dei tuoi elementi `.card` deve usare flexbox per giustificare i suoi figli usando `space-between`.
1. Ciascuno dei tuoi elementi `.left` deve usare proprietà flexbox degli elementi per allinearsi all'inizio del suo contenitore genitore.
1. Ciascuno dei tuoi elementi `.middle` deve usare proprietà flexbox degli elementi per allinearsi al centro del suo contenitore genitore.
1. Ciascuno dei tuoi elementi `.right` deve usare proprietà flexbox degli elementi per allinearsi alla fine del suo contenitore genitore.
1. Ciascuno dei tuoi elementi `.middle` deve usare flexbox per mostrare i suoi figli in colonna.

Ecco alcuni caratteri che puoi copiare e incollare per usare nelle tue carte se vuoi: `♠`, `♣`, `♥` e `♦`.

# --hints--

Devi avere un elemento `main` con un `id` di `playing-cards`.

```js
assert.exists(document.querySelector('main#playing-cards'));
```

Devi avere almeno tre elementi `div` con una classe `card` all'interno del tuo elemento `#playing-cards`.

```js
assert.isAtLeast(document.querySelectorAll('#playing-cards > div.card')?.length, 3);
```

Ciascuno dei tuoi elementi `.card` deve avere una larghezza e un'altezza.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  assert.isAbove(el.clientWidth, 0);
  assert.isAbove(el.clientHeight, 0);
});
```

Ciascuno dei tuoi elementi `.card` deve avere esattamente tre elementi `div` come figli.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  const children = el.children;
  assert.lengthOf(children, 3);
  Array.from(children).forEach(child => {
    assert.equal(child.tagName, 'DIV');
  });
});
```

Devi avere un `div` con una classe `left` all'interno di ogni elemento `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.left')));
```

Devi avere un `div` con una classe `middle` all'interno di ogni elemento `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.middle')));
```

Devi avere un `div` con una classe `right` all'interno di ogni elemento `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.right')));
```

Devi avere un selettore `#playing-cards` che imposta il `display` dei suoi elementi su `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.display, 'flex');
```

Devi avere un selettore `#playing-cards` che imposta il `justify-content` dei suoi elementi su `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.justifyContent, 'center');
```

Devi avere un selettore `#playing-cards` che imposta il `flex-wrap` dei suoi elementi su `wrap`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.flexWrap, 'wrap');
```

Devi avere un selettore `#playing-cards` che imposta il `gap` dei suoi elementi su `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.gap, '20px');
```

Devi avere un selettore `.card` che imposta il `display` dei suoi elementi su `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.display, 'flex');
```

Devi avere un selettore `.card` che imposta il `justify-content` dei suoi elementi su `space-between`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.justifyContent, 'space-between');
```

Devi avere un selettore `.left` che imposta il `align-self` dei suoi elementi su `flex-start`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left')?.alignSelf, 'flex-start');
```

Devi avere un selettore `.middle` che imposta il `align-self` dei suoi elementi su `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.alignSelf, 'center');
```

Devi avere un selettore `.right` che imposta il `align-self` dei suoi elementi su `flex-end`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.right')?.alignSelf, 'flex-end');
```

Devi avere un selettore `.middle` che imposta il `display` dei suoi elementi su `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.display, 'flex');
```

Devi avere un `flex-direction` impostato su `column` nella tua classe `.middle`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Playing Cards</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Straight</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="playing-cards">
    <div class="card">
      <div class="left">
        <div>A</div>
        <div>♠</div>
      </div>
      <div class="middle">
        <div>♠</div>
      </div>
      <div class="right">
        <div class="flipped">♠</div>
        <div class="flipped">A</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>2</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div>♣</div>
        <div class="flipped">♣</div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">2</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>3</div>
        <div>♥</div>
      </div>
      <div class="middle">
        <div>♥</div>
        <div>♥</div>
        <div class="flipped">♥</div>
      </div>
      <div class="right">
        <div class="flipped">♥</div>
        <div class="flipped">3</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>4</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♣</div>
          <div>♣</div>
        </div>
        <div class="middle-section">
          <div class="flipped">♣</div>
          <div class="flipped">♣</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">4</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>5</div>
        <div>♦</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♦</div>
        <div class="flipped">5</div>
      </div>
    </div>
  </main>
</body>

</html>
```

```css
body {
  background-color: Thistle;
  min-height: 100vh;
}

.red {
  color: red;
}

.flipped {
  transform: rotate(180deg);
}

#playing-cards {
  margin-top: 100px;
  display: flex;
  justify-content: center;
  gap: 20px;
  flex-wrap: wrap;
}

.card {
  background-color: white;
  border: 1px solid black;
  border-radius: 10px;
  box-shadow: 2px 2px 2px gray;
  min-width: 250px;
  max-width: 250px;
  height: 350px;
  font-size: 30px;
  font-weight: bold;
  padding: 5px;
  box-sizing: border-box;
  display: flex;
  justify-content: space-between;
  text-align: center;
}

.left {
  align-self: flex-start;
}

.middle {
  align-self: center;
  display: flex;
  font-size: 80px;
  flex-direction: column;
}

.right {
  align-self: flex-end;
}

.middle-section {
  display: flex;
  justify-content: center;
}
```
