---
id: ac6993d51946422351508a41
title: 实现一个字符串截断算法
challengeType: 26
dashedName: truncate-a-string
---

# --description--

在本实验中，你将练习将字符串截断到一定长度。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个接受两个参数的函数 `truncateString`，第一个参数是字符串，第二个参数是数字。
1. 如果字符串的长度大于给定的数字，则应截断字符串以减少长度，使其等于给定的数字，并在截断字符串的末尾添加 `...`。
1. 如果字符串长度等于或小于给定的数字，则返回的字符串应保持不变。

# --hints--

`truncateString("A-tisket a-tasket A green and yellow basket", 8)` 应该返回字符串 `A-tisket...`。

```js
assert.strictEqual(
  truncateString('A-tisket a-tasket A green and yellow basket', 8),
  'A-tisket...'
);
```

`truncateString("Peter Piper picked a peck of pickled peppers", 11)` 应该返回字符串 `Peter Piper...`。

```js
assert.strictEqual(
  truncateString('Peter Piper picked a peck of pickled peppers', 11),
  'Peter Piper...'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length)` 应该返回字符串 `A-tisket a-tasket A green and yellow basket`。

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length + 2)` 应该返回字符串 `A-tisket a-tasket A green and yellow basket`。

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length + 2
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-", 1)` 应该返回字符串 `A...`。

```js
assert.strictEqual(truncateString('A-', 1), 'A...');
```

`truncateString("Absolutely Longer", 2)` 应该返回字符串 `Ab...`。

```js
assert.strictEqual(truncateString('Absolutely Longer', 2), 'Ab...');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function truncateString(str, num) {
  if (num >= str.length) {
    return str;
  }

  return str.slice(0, num) + '...';
}

```
