---
id: 694ce8609fb805ba21cebc02
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

تحتاج وحدة التحكم المهمة إلى طريقة لتبديل مواضع عضوين من الطاقم في القائمة. أنشئ دالة فارغة باسم `swapCrewMembers` بثلاث معلمات:

- `crew`: مصفوفة من كائنات رواد الفضاء

- `fromIndex`: فهرس رائد الفضاء الأول الذي سيتم تبديله

- `toIndex`: فهرس رائد الفضاء الثاني الذي سيتم تبديله

# --hints--

يجب أن تنشئ دالة باسم `swapCrewMembers`.

```js
assert.isFunction(swapCrewMembers);
```

يجب أن تحتوي دالة `swapCrewMembers` على المعلمات `crew` و `fromIndex` و `toIndex`.

```js
const regex = __helpers.functionRegex('swapCrewMembers', ['crew', 'fromIndex', 'toIndex']);
assert.match(__helpers.removeJSComments(swapCrewMembers.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

--fcc-editable-region--

--fcc-editable-region--
```
