---
id: 694c40d86fe6a78d20b6559c
title: خطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستستفيد من JavaScript لبناء قائمة رواد فضاء لمهمة فضائية.

ابدأ بإنشاء مصفوفة فارغة لقائمتك باسم `squad`.

# --hints--

يجب أن تنشئ متغيرًا باسم `squad`.

```js
assert.isDefined(squad);
```

يجب أن يكون المتغير `squad` مصفوفة.

```js
assert.isArray(squad);
```

يجب أن تكون المصفوفة `squad` فارغة.

```js
assert.deepEqual(squad, []);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
