---
id: 694c84106e9023f289c03b84
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

تبدأ قائمتك برائد فضاء واحد بدور القائد. أنشئ كائنًا باسم `firstAstronaut` بالخصائص التالية:

| المفتاح | القيمة |
| ----------- | ------- |
|`id`|`1`|
|`name`|`"Andy"`|
|`role`|`"Commander"`|
|`isEVAEligible`|`true`|
|`priority`|`3`|

# --hints--

يجب أن تنشئ متغيرًا باسم `firstAstronaut`.

```js
assert.isDefined(firstAstronaut);
```

يجب أن يكون `firstAstronaut` كائنًا.

```js
assert.isObject(firstAstronaut);
```

يجب أن يتضمن كائن `firstAstronaut` القيم الصحيحة للخصائص.

```js
assert.deepEqual(firstAstronaut, { id: 1, name: "Andy", role: "Commander", isEVAEligible: true, priority: 3 });
```

# --seed--

## --seed-contents--

```js
const squad = [];

--fcc-editable-region--

--fcc-editable-region--
```
