---
id: 694c9c186c089f7bdff0f3c6
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

بشكل عام، من الممارسات الجيدة التحقق من صحة المُدخَلات داخل الدالة. استخدم حلقة `for` للتكرار خلال مصفوفة `crew` وتحقق مما إذا كان أي عضو لديه نفس `id` مثل `astronaut` المُدخَل. إذا وُجد تكرار، سجّل رسالة الخطأ `console.log("Duplicate ID: " + astronaut.id)` ثم استدعِ `return` للخروج المبكر من الدالة.

إليك مثالًا على هذه التقنية المستخدمة بشكل شائع:

```js
if (existingMember.id === newMember.id) {
    console.log("Duplicate ID: " + newMember.id);
    return;
}
```

# --hints--

يجب أن تستخدم دالتك حلقة `for` لمسح أعضاء `crew` الحاليين (أي صيغة صحيحة لـ `for (...)` مقبولة).

```js
assert.match(__helpers.removeJSComments(code), /for\s*\([^)]*\)/);
```

داخل حلقة `for`، تحقق من وجود `id` مكرر. إذا وُجد، سجّل `"Duplicate ID: " + astronaut.id"` في وحدة التحكم ثم استدعِ `return` للخروج من الدالة.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewWithDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 },
    { id: 3 }
  ];
  addCrewMember(crewWithDuplicate, { id: 3 });

  assert.lengthOf(
    spy.calls,
    1,
    "Duplicate ID input should trigger exactly one console.log and then exit"
  );
  assert.match(
    spy.calls[0][0],
    /duplicate\s+id\s*:\s*3/i,
    "Console message should mention 'duplicate id'"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

عند إعطاء مُدخَلات صحيحة، يجب ألا تستدعي دالة `addCrewMember` رسالة خطأ في `console.log()` أو خروجًا مبكرًا بـ `return`.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewNoDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 }
  ];
  addCrewMember(crewNoDuplicate, { id: 5 });

  assert.lengthOf(
    spy.calls,
    0,
    "A truly new astronaut should produce zero console.log calls"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
--fcc-editable-region--

--fcc-editable-region--
}
```
