---
id: 694c84330cbe604821b9deef
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

يتطلب مركز التحكم طريقة لإضافة أعضاء جدد إلى القائمة. أنشئ دالة فارغة باسم `addCrewMember` تستقبل معلمتين، `crew` و `astronaut`.

# --hints--

يجب أن تنشئ دالة باسم `addCrewMember`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addCrewMember);
```

يجب أن تحتوي دالتك `addCrewMember` على معلمتين هما `crew` و `astronaut`.

```js
const explorer = await __helpers.Explorer(code);
const { addCrewMember } = explorer.allFunctions;
assert.lengthOf(addCrewMember?.parameters, 2);
assert.equal(addCrewMember?.parameters?.[0]?.toString(), "crew");
assert.equal(addCrewMember?.parameters?.[1]?.toString(), "astronaut");
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
