---
id: 694d8ddd106340237efb0c80
title: الخطوة 12
challengeType: 1
dashedName: step-12
---

# --description--

تستطيع طريقة `splice()` تعديل المصفوفات بإضافة أو إزالة عناصر في أي موضع، بما في ذلك المنتصف. وبما أن `splice()` تُرجع مصفوفة تحتوي على العناصر المُزالة، يمكنك استخدامها لتبديل عنصرين في مصفوفة دون تغيير الأصل. يمكن تنفيذ ذلك في سطر واحد باستخدام التقنية التالية:

```js
// swap elements at i and j without mutating the original
const copy = array.slice();
copy[i] = copy.splice(j, 1, copy[i])[0];
```

إليك مثال:

```js
const originalArray = [12, 97, 68, 55];
const copyArray = originalArray.slice();
copyArray[1] = copyArray.splice(3, 1, copyArray[1])[0];
console.log(copyArray); // [12, 55, 68, 97]
```

التقنية، المطبقة في السطر الثالث، تعمل كما يلي:

- `splice(3, 1, copyArray[1])` تزيل العنصر عند الفهرس `3` (`55`)

- تُدرج العنصر من `copyArray[1]` (`97`) في الفهرس `3`

- تُرجع `splice()` المصفوفة `[55]` التي تحتوي على العنصر المُزال

- يستخرج `[0]` القيمة `55` من تلك المصفوفة

- تُعيَّن تلك القيمة (`55`) مرة أخرى إلى `copyArray[1]`، مكتملةً عملية التبديل

استخدم هذه التقنية مع مصفوفة `updatedCrew` لتبديل رواد الفضاء عند `fromIndex` و `toIndex`.

# --hints--

يجب أن تنفذ التبديل باستدعاء طريقة `splice()`.

```js
assert.match(swapCrewMembers.toString(), /splice\s*\(/);
```

يجب استدعاء `splice()` مع مصفوفة `updatedCrew` لتنفيذ التبديل. استخدم التقنية ذات السطر الواحد.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(swapCrewMembers.toString()));
const regex = /updatedCrew\[(fromIndex|toIndex)\]=updatedCrew\.splice\((fromIndex|toIndex),1,updatedCrew\[\1\]\)\[0\];?\s*\}$/;
assert.match(cleaned, regex);
```

يجب ألا تُغيّر دالة `swapCrewMembers` مصفوفة `crew` الأصلية المُدخلة.

```js
const demo = [{ id: 1 }, { id: 2 }, { id: 3 }];
swapCrewMembers(demo, 0, 2);
assert.deepEqual(demo.map(a => a.id), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
  --fcc-editable-region--
  
  --fcc-editable-region--
}
```
