---
id: 694e1931f6509ae16aa1c36f
title: الخطوة 18
challengeType: 1
dashedName: step-18
---

# --description--

أبلغك مركز التحكم أن قائمة رواد الفضاء المؤهلين لـ EVA يجب أن تُرتب أيضًا حسب الأولوية تنازليًا. هناك عدة طرق لترتيب المصفوفة - وربما أبسطها هو `bubble sort`.

تعمل طريقة `Bubble sort` من خلال المرور المتكرر على القائمة، ومقارنة العناصر المجاورة، وتبديلها إذا كانت في ترتيب خاطئ. بعد كل مرور، يتحرك العنصر الذي يجب أن يأتي أولًا بناءً على معايير الفرز الخاصة بك أقرب إلى (أو "يطفو" نحو) موقعه الصحيح في المصفوفة. إليك كيف يمكنك فرز `crew` حسب الأولوية تنازليًا باستخدام `bubble sort`:

```js
// Outer loop: controls how many passes we make
for (let i = 0; i < crew.length - 1; i++) {
  // Inner loop: compares neighboring items
  for (let j = 0; j < crew.length - 1 - i; j++) {
    // If current member has lower priority than next, swap
    if (crew[j].priority < crew[j + 1].priority) {
      // Using a temp variable for the swap
        const temp = crew[j];
        crew[j] = crew[j + 1];
        crew[j + 1] = temp;
    }
  }
}
```

أنشئ دالة مساعدة جديدة باسم `sortByPriorityDescending` تستقبل معلمة `crew`. يجب أن ترتب هذه الدالة مباشرة مصفوفة `crew` المُدخلة حسب الأولوية تنازليًا، باستخدام حلقتي `for` متداخلتين، ولا تُرجع أي قيمة.

# --hints--

يجب أن يكون لديك دالة باسم `sortByPriorityDescending` مع معلمة `crew`.

```js
assert.isFunction(sortByPriorityDescending);
const regex = __helpers.functionRegex('sortByPriorityDescending', ['crew']);
const cleaned = __helpers.removeJSComments(sortByPriorityDescending.toString());
assert.match(cleaned, regex);
```

يجب أن تستخدم دالة `sortByPriorityDescending` حلقتي `for` متداخلتين لتنفيذ الترتيب.

```js
const loopMatches = sortByPriorityDescending.toString().match(/for\s*\(/g) || [];
assert.isAtLeast(loopMatches.length, 2);
```

يجب أن ترتب دالة `sortByPriorityDescending` مصفوفة `crew` المُدخلة حسب `priority` بترتيب تنازلي.

```js
const crewSample = [
  { name: "A", priority: 2 },
  { name: "B", priority: 5 },
  { name: "C", priority: 3 }
];
sortByPriorityDescending(crewSample);
assert.deepEqual(crewSample.map(a => a.name), ["B", "C", "A"]);
```

يجب أن ترتب دالة `sortByPriorityDescending` مصفوفة `crew` المُدخلة مباشرة في مكانها دون إرجاع مصفوفة جديدة.

```js
const crewSample2 = [
  { name: "X", priority: 1 },
  { name: "Y", priority: 2 }
];
const result = sortByPriorityDescending(crewSample2);
assert.strictEqual(result, undefined);
assert.deepEqual(crewSample2.map(a => a.priority), [2, 1]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
  updatedCrew[fromIndex] = updatedCrew.splice(toIndex, 1, updatedCrew[fromIndex])[0];

  return updatedCrew; 
}

const updatedSquad = swapCrewMembers(squad, 2, 5);

--fcc-editable-region--

--fcc-editable-region--

function getEVAReadyCrew(crew) {
  const eligible = [];
  for (const astronaut of crew) {
    if (astronaut.isEVAEligible) eligible.push(astronaut);
  }
  
  return eligible;
}
```
