---
id: 68b94f23f7b3ee14078c89c5
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utatumia kipengele cha `iframe` kuonyesha video. Muundo wa msingi wa HTML umeandaliwa kwa ajili yako.

Anza kwa kuunda kipengele cha `h1` chenye maandishi `iframe Video Display`.

# --hints--

Unapaswa kuwa na kipengele cha `h1`.

```js
assert.exists(document.querySelector('h1'))
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi `iframe Video Display`.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), 'iframe Video Display')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
