---
id: 68ba9c4f1688914d72876097
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Ongeza `encrypted-media`, `gyroscope`, na `web-share` kwa thamani zilizopo katika sifa ya kipengele ya `allow`.

Hizi tatu zitaruhusu matumizi ya nyongeza za vyombo vya habari vilivyohifadhiwa kwa usimbaji ili kulinda video, kutoa upatikanaji wa sensa za mwendo na mwelekeo wa kifaa, na kuruhusu kushiriki maudhui ya `iframe` kupitia madirisha ya kushiriki ya asili ya kifaa.

# --hints--

Sifa ya kipengele ya `allow` ya kipengele chako cha `iframe` inapaswa kuwa na `encrypted-media` kama mojawapo ya thamani zake.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /encrypted-media/)
```

Sifa ya kipengele ya `allow` ya kipengele chako cha `iframe` inapaswa kuwa na `gyroscope` kama mojawapo ya thamani zake.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /gyroscope/)
```

Sifa ya kipengele ya `allow` ya kipengele chako cha `iframe` inapaswa kuwa na `web-share` kama mojawapo ya thamani zake.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /web-share/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      allow="accelerometer autoplay clipboard-write"
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
