---
id: 68ba9c4f1688914d72876096
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Moja ya sifa za kipengele ni `allow`. Ni kama orodha ya ruhusa inayomwambia kivinjari ni vipengele gani `iframe` inaruhusiwa kutumia.

Hapa kuna kipengele cha `iframe` chenye sifa ya kipengele `allow`:

```html
<iframe
  allow="accelerometer autoplay clipboard-write encrypted-media gyroscope picture-in-picture web-share"
></iframe>
```

Ongeza sifa ya kipengele `allow` yenye thamani `accelerometer`, `autoplay`, na `clipboard-write`.

`accelerometer` inaruhusu `iframe` kutumia sensa za mwendo ili itambue mambo kama kuinua na kuzungusha kifaa. `autoplay` inaruhusu video kuanza kucheza moja kwa moja, na `clipboard-write` inaruhusu iframe kuandika data kwenye ubao wa kunakili wa mtumizi.

# --hints--

Kipengele chako cha `iframe` kinapaswa kuwa na sifa ya kipengele `allow` iliyowekwa kuwa `accelerometer`.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /accelerometer/)
```

Sifa ya kipengele `allow` ya kipengele chako cha `iframe` inapaswa kuwa na `autoplay` kama mojawapo ya thamani zake.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /autoplay/)
```

Sifa ya kipengele `allow` ya kipengele chako cha `iframe` inapaswa kuwa na `clipboard-write` kama mojawapo ya thamani zake.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /clipboard-write/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
