---
id: 68ba9c4f1688914d72876095
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Kipengele cha `iframe` pia kinachukua sifa ya kipengele ya `src` yenye thamani inayobainisha URL au njia ya rasilimali ya kuonyeshwa.

Ongeza sifa ya kipengele ya `src` yenye thamani ya `https://www.youtube.com/embed/I0_951_MPE0` kwenye kipengele chako cha `iframe`.

Kwa sasa, unapaswa kuona video ikionyeshwa kwenye ukurasa, lakini bado kuna sifa za kipengele zaidi unazohitaji kuongeza.

# --hints--

Kipengele chako cha `iframe` kinapaswa kuwa na sifa ya kipengele ya `src` iliyowekwa kuwa `https://www.youtube.com/embed/I0_951_MPE0`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('src'), 'https://www.youtube.com/embed/I0_951_MPE0')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
