---
id: 68ba9c4f1688914d72876098
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Sifa ya kipengele inayofuata utakayoongeza ni `referrerpolicy`. Hii ni sheria inayobainisha ni maelezo gani unayoshiriki unapounganisha ukurasa wako na ukurasa mwingine.

Ongeza sifa ya kipengele `referrerpolicy` na uiweke kuwa `strict-origin-when-cross-origin`. Hii inashiriki anwani kamili kwenye tovuti ile ile, jina la tovuti tu kwenye tovuti nyingine, na haina chochote kwenye tovuti zisizo salama.

# --hints--

Kipengele chako cha `iframe` kinapaswa kuwa na sifa ya kipengele `referrerpolicy` iliyowekwa kuwa `strict-origin-when-cross-origin`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('referrerpolicy'), 'strict-origin-when-cross-origin')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
      allow="accelerometer autoplay clipboard-write encrypted-media gyroscope web-share"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
