---
id: 68ba9c4f1688914d72876094
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Katika funzo la kwanza kuhusu kipengele cha `iframe`, ulijifunza kuwa ni kipengele kilichobadilishwa kama vile `img`. Hii inamaanisha pia kinaweza kuchukua sifa za kipengele za `width` na `height` ili kuamua urefu na upana wake.

Toa kipengele chako cha `iframe` sifa ya `width` ya `560` na sifa ya `height` ya `315`.

# --hints--

Kipengele chako cha `iframe` kinapaswa kuwa na sifa ya `width` iliyowekwa kuwa `560`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('width'), '560')
```

Kipengele chako cha `iframe` kinapaswa kuwa na sifa ya `height` iliyowekwa kuwa `315`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('height'), '315')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>

    <iframe
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
