---
id: 687d7deea6a2a22fa8816a73
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在之前的课程中，你学习了一些常见的数据结构。在本次研讨会中，你将实现一个链表。

首先，创建一个名为 `LinkedList` 的类来表现将保存你的列表节点的结构。在类内部添加 `pass` 关键字，以便测试能够通过。

# --hints--

你应该创建一个名为 `LinkedList` 的类。记得添加 `pass` 关键字。

```js
({
    test: () => runPython(`assert _Node(_code).has_class("LinkedList")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
