---
id: 687d8e05a529a4838f1c961c
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

接下来，删除 `Node` 类中的 `pass` 关键字，然后为 `Node` 类添加一个带有两个参数 `self` 和 `element` 的 `__init__` 方法。

在该方法内，将 `element` 赋值给 `self.element`。这会保存创建节点时传入的值。

# --hints--

你应该在你的 `Node` 类中定义一个 `__init__` 方法。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").has_function("__init__")`) 
})
```

你的 `__init__` 方法应有两个参数，`self` 和 `element`。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").find_function("__init__").has_args("self, element")`) 
})
```

你应该在 `__init__` 方法内将 `element` 赋值给 `self.element`。

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.element == "A"
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
--fcc-editable-region--
    class Node:
        pass
--fcc-editable-region--
```
