---
id: 687d8c6f066877768f7ee441
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

在你的 `LinkedList` 类中，将 `pass` 关键字替换为一个名为 `Node` 的内嵌类。它将表现保存列表中数据的单个节点。

# --hints--

你应该在 `LinkedList` 类中创建一个 `Node` 类。记得在 `Node` 类中添加 `pass` 关键字。

```js
({
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").has_class("Node")`) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class LinkedList:
    pass
--fcc-editable-region--
```
