---
id: 687e344cc1993f7316af030c
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

在 `LinkedList` 类中，添加一个 `__init__` 方法，将 `self.length` 设置为 `0`。这会初始化一个计数器，用于跟踪列表中当前有多少个节点（元素）。

# --hints--

你应该在 `LinkedList` 类的 `__init__` 方法中将 `self.length` 设置为 `0`。

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.length == 0`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    
--fcc-editable-region--
```
