---
id: 68800a64b06e80d9a31c131b
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

仍然在 `LinkedList` 类的 `__init__` 方法中，将 `self.head` 设置为 `None`。这作为列表中第一个节点的引用。此时列表为空，因此你必须将其设置为 `None`。


# --hints--

你应该在 `LinkedList` 类的 `__init__` 方法中将 `self.head` 设置为 `None`。

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.head is None`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    def __init__(self):
        self.length = 0
--fcc-editable-region--
```
