---
id: 688003f589b875af60273dd9
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

仍然在 `__init__` 方法内，赋值 `None` 给 `self.next`。这将 `next` 指针设置为 `None`，意味着新节点尚未指向列表中的任何其他节点。

稍后，`next` 可以更新为指向另一个节点，以形成链表的链。

# --hints--

在 `__init__` 方法中，你应该将 `next` 属性初始化为 `None`。

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.next is None
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
--fcc-editable-region--
        def __init__(self, element):
            self.element = element
--fcc-editable-region--
```
