---
id: 68857e0a49baf4751f737a2a
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

现在，在 `is_empty` 方法之后，为 `LinkedList` 类创建一个名为 `add` 的方法。新方法应有两个参数，`self` 和 `element`。

在你的新方法中，创建一个名为 `node` 的变量，并通过使用 `self.Node(element)` 为它赋值一个新的 `Node` 实例。这将创建一个要添加到链表中的新节点。

# --hints--

你应该在 `LinkedList` 类中创建一个 `add` 方法。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").has_function("add")`
      )
    )
})
```

你的 `add` 方法应有两个参数，`self` 和 `element`。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("add").has_args("self, element")`) 
})
```

你应该在 `add` 方法中将 `self.Node(element)` 赋值给 `node`。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_body().is_equivalent("node = self.Node(element)")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
