---
id: 68858a1b4e867acd87f6f7e8
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

如果列表为空，则新节点成为第一个元素（`head`）。在 `add` 方法中，创建一个 `if` 语句，使用 `self.is_empty()` 查看列表是否为空，并将 `node` 赋值给 `self.head`。

# --hints--

你应该使用 `is_empty()` 方法查看列表是否为空。

```js
({
  test: () => runPython(`
  _cond = _Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_conditions()[0]
  assert _cond.is_equivalent("self.is_empty()")
`)})
```

你应该在 `if` 语句内将 `node` 赋值给 `self.head`。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_body().is_equivalent("self.head = node")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
