---
id: 68868f502659139bed8e5c90
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

在为列表添加新建节点后，你需要更新它的长度。在 `add` 方法的末尾，将 `self.length` 增加 `1`。

# --hints--

你应该在 `add` 方法的末尾将 `self.length` 增加 `1`。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").has_stmt("self.length += 1")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
