---
id: 68867c3898b09115d8c5140f
title: 步骤 13
challengeType: 20
dashedName: step-13
---

# --description--

你刚创建的 `while` 循环使你能够遍历列表中的每个节点，直到到达最后一个节点。一旦发生这种情况，你可以将节点添加到列表中。

在 `while` 块之后、`else` 块中，赋值 `node` 为 `current_node.next`。

# --hints--

你应该在 `else` 块中将 `node` 赋值给 `current_node.next`。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].has_stmt("current_node.next = node")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
