---
id: 6888d05ba074461b64f8eaca
title: 步骤 16
challengeType: 20
dashedName: step-16
---

# --description--

接下来，你将创建另一个方法，用于从链表中移除给定的元素。

创建一个名为 `remove` 的方法，带有两个参数，`self` 和 `element`。在其中，创建一个名为 `previous_node` 的变量并赋值为 `None`。

# --hints--

你应该创建一个名为 `remove` 的方法。

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_class("LinkedList").has_function("remove")
    `)
  )
})
```

你的 `remove` 方法应有两个参数，`self` 和 `element`。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("remove").has_args("self, element")`) 
})
```

你应该在 `remove` 方法中将 `None` 赋值给 `previous_node`。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("remove").find_body().is_equivalent("previous_node = None")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
