---
id: 6888f9827637322d4f309192
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

在 `remove` 方法内，创建一个名为 `current_node` 的变量并赋值为 `self.head`。

这将从第一个节点开始遍历链表。`current_node` 将用于跟踪当前正在检查的节点，以便查找包含指定元素以删除的节点。

# --hints--

你应该在 `remove` 方法内将 `self.head` 赋值给 `current_node`。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("remove").has_stmt("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

--fcc-editable-region--
    def remove(self, element):
        previous_node = None
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
