---
id: 688acb2053ac7978a06345f1
title: 步骤21
challengeType: 20
dashedName: step-21
---

# --description--

接下来，创建一个 `elif` 语句，查看 `previous_node` 是否不为 `None`。在其中将 `current_node.next` 赋值为 `previous_node.next`。

这样，如果找到要删除的元素且它不在头节点，前一个节点的 `next` 指针将被更新以跳过当前节点（要删除的元素），从链表中移除它。

# --hints--

你应该创建一个 `elif` 语句来查看 `previous_node` 不是 `None`。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("previous_node is not None")
`) })
```

你应该在 `elif` 语句内将 `current_node.next` 赋值给 `previous_node.next`。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0]
assert _cond.is_equivalent("if current_node is None:\\n return\\nelif previous_node is not None:\\n previous_node.next = current_node.next")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
        while current_node is not None and current_node.element != element:
            previous_node = current_node
            current_node = current_node.next
--fcc-editable-region--
        if current_node is None:
            return
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
