---
id: 68b84df2648b4c96b25022f3
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

让我们测试你到目前为止的链表！在你的代码底部，创建一个名为 `my_list` 的 `LinkedList` 实例。然后，使用你在上一步创建的 `is_empty` 方法，打印 `my_list.is_empty()`。

# --hints--

你应该创建一个名为 `my_list` 的 `LinkedList` 实例。

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_variable("my_list").value_is_call("LinkedList")
    `)
  )
})
```

你应该打印 `my_list.is_empty()`。

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).has_call("print(my_list.is_empty())")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0

--fcc-editable-region--

--fcc-editable-region--
```
