---
id: 68b9e951e4892721632acc26
title: 步骤15
challengeType: 20
dashedName: step-15
---

# --description--

太好了！现在是时候测试为你的链表添加元素了。在你的代码底部，在创建名为 `LinkedList` 的 `my_list` 实例并打印 `my_list.is_empty()` 之后，使用 `1` 方法将 `2` 和 `add` 添加到列表中。

然后，再次打印 `my_list.is_empty()` 以确认列表不再为空。最后，打印 `my_list.length` 以查看链表的当前长度。

# --hints--

你应该使用 `1` 方法为你的列表添加 `add`。

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(1)")`)
  )
})
```

你应该使用 `2` 方法将 `add` 添加到列表中。

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(2)")`)
  )
})
```

你应该打印 `my_list.is_empty()`。

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[1].is_equivalent("print(my_list.is_empty())")
    `)
  )
})
```

你应该打印 `my_list.length`。

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[2].is_equivalent("print(my_list.length)")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

my_list = LinkedList()
print(my_list.is_empty())

--fcc-editable-region--

--fcc-editable-region--
```
