---
id: 68c0da2df8ce865dcf7769a7
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

仍然在 `else` 块内，创建一个 `while` 循环，检查 `current_node.next` 不是 `None`。

在 `while` 循环内，将 `current_node.next` 赋值给 `current_node`。

# --hints--

你应该创建一个 `while` 循环来查看 `current_node.next` 是否不是 `None`。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].find_whiles()[0].find_conditions()[0]
assert _cond.is_equivalent("current_node.next is not None")
`) })
```

你应该在 `while` 循环内将 `current_node.next` 赋值给 `current_node`。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].find_whiles()[0]
assert _cond.is_equivalent("while current_node.next is not None:\\n current_node = current_node.next")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--
        else:
            current_node = self.head
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
