---
id: 60a3e3396c7b40068ad6996b
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Within the `head` element, add a `meta` tag which sets the `charset` to `UTF-8`, and a `title` element with the value `Rothko Painting`.

Dentro del elemento `body`, agregue un elemento `img` con un `src` de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.
# --hints--

El código debe tener una etiqueta `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

La etiqueta `meta` debe establecer el `charset` en `UTF-8`.

```js
assert(document.querySelector('meta').getAttribute('charset')?.toLowerCase() === 'utf-8');
```

El código debe tener un elemento `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

El `title` debe ser `Rothko Painting`.

```js
assert(document.querySelector('title').innerText === 'Rothko Painting');
```

El código debe tener un elemento `img`.

```js
assert(document.querySelectorAll('img').length === 1);
```

El elemento `img` debe tener un `src` de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
