---
id: 60a3e3396c7b40068ad6996e
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Margin is the area outside of the box, and can be used to control the space between other boxes or elements.

Aquí el elemento inferior tiene un margen superior más grande, empujándolo más abajo en la página.

Ahora que entiendes el modelo de caja CSS, comencemos con la pintura de Rothko.

Elimine el elemento`<img>`.

# --hints--

No debe tener un elemento `img` en el código.

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
