---
id: bad87fee1348bd9bec908846
title: ブートストラップの行を作成する
challengeType: 0
forumTopicId: 16813
dashedName: create-a-bootstrap-row
---

# --description--

これからインライン要素のためにブートストラップの行を作成します。

`div` タグの下に `h3` 要素を作成し、クラスに `row` を指定してください。

# --hints--

`div` 要素の下に `h3` 要素を追加してください。

```js
assert.lengthOf(document.querySelectorAll('div'),2);
assert.lengthOf(document.querySelectorAll('div.row h3.text-primary'),0);
assert.lengthOf(document.querySelectorAll('div.row + h3.text-primary'),0)
assert.lengthOf(document.querySelectorAll('h3.text-primary + div.row'),1);
```

`div` 要素にはクラス `row` を指定してください。

```js
const newDiv = document.querySelectorAll('div')?.[1]; 
assert.isTrue(newDiv?.classList?.contains('row'));
```

`row div` は `container-fluid div` の中にネストしてください。

```js
assert.lengthOf(document.querySelectorAll('div.container-fluid div.row'),1);
```

`div` 要素には終了タグが必要です。

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g).length ,code.match(/<div/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>

</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row"></div>
</div>
```
