---
id: bad87fee1348bd9aec908847
title: ブートストラップのRowを分割する
challengeType: 0
forumTopicId: 18306
dashedName: split-your-bootstrap-row
---

# --description--

ブートストラップのRowができたので、要素を収めるために2つのカラムに分割しましょう。

Rowの中に`div`要素を2つ作成し、両方にクラス`col-xs-6`を付けてください。

# --hints--

`div class="col-xs-6"`要素を2つ、`div class="row"`要素の中にネストしてください。

```js
assert.lengthOf(document.querySelectorAll('div.row > div.col-xs-6'),2);
```

すべての`div`要素には終了タグが必要です。

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g).length , code.match(/<div/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">


  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6"></div>
    <div class="col-xs-6"></div>
  </div>
</div>
```
