---
id: bad87fee1348bd9aec908848
title: ブートストラップのウェルを作成する
challengeType: 0
forumTopicId: 16825
dashedName: create-bootstrap-wells
---

# --description--

Bootstrapには、カラムに奥行きの視覚効果を作り出すクラス`well`があります。

各`div` `well`要素の中に、クラス`col-xs-6`を持つ`div`要素をネストしてください。

# --hints--

各クラス`div`を持つ`well`要素の中に、クラス`div`を持つ`col-xs-6`要素を追加してください。

```js
const wells = document.querySelectorAll('div.col-xs-6 > div.well');
assert.lengthOf( wells,2 ); 
```

クラス`div`を持つ`col-xs-6`要素の中に、クラス`div`を持つ`row`要素が両方ともネストされている必要があります。

```js
assert.lengthOf(document.querySelectorAll('div.row > div.col-xs-6'),2);
```

すべての`div`要素には終了タグが必要です。

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g)?.length , code.match(/<div/g)?.length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">

    </div>
    <div class="col-xs-6">

    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well"></div>
    </div>
    <div class="col-xs-6">
      <div class="well"></div>
    </div>
  </div>
</div>
```
