---
id: bad87fee1348bd9aec908856
title: ブートストラップボタンにラベルを付ける
challengeType: 0
forumTopicId: 18222
dashedName: label-bootstrap-buttons
---

# --description--

井戸にラベルを付けたのと同じように、ボタンにもラベルを付けたいです。

それぞれの`button`要素に、そのidセレクターに対応するテキストを付けてください。

# --hints--

idが`target1`の`button`要素には、テキスト`#target1`を付けてください。

```js
assert.match(document.querySelector('#target1')?.textContent,/#target1/gi)
```

idが`target2`の`button`要素には、テキスト`#target2`を付けてください。

```js
assert.match(document.querySelector('#target2')?.textContent,/#target2/gi)
```

idが`target3`の`button`要素には、テキスト`#target3`を付けてください。

```js
assert.match(document.querySelector('#target3')?.textContent,/#target3/gi)
```

idが`target4`の`button`要素には、テキスト`#target4`を付けてください。

```js
assert.match(document.querySelector('#target4')?.textContent,/#target4/gi)
```

idが`target5`の`button`要素には、テキスト`#target5`を付けてください。

```js
assert.match(document.querySelector('#target5')?.textContent,/#target5/gi)
```

idが`target6`の`button`要素には、テキスト`#target6`を付けてください。

```js
assert.match(document.querySelector('#target6')?.textContent,/#target6/gi)
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1"></button>
        <button class="btn btn-default target" id="target2"></button>
        <button class="btn btn-default target" id="target3"></button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4"></button>
        <button class="btn btn-default target" id="target5"></button>
        <button class="btn btn-default target" id="target6"></button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
