---
id: bad87fee1348bd9aec908846
title: ブートストラップの見出しを作成する
challengeType: 0
forumTopicId: 16812
dashedName: create-a-bootstrap-headline
---

# --description--

それでは、HTML、CSS、ブートストラップのスキルを練習するために、ゼロから何かを作ってみましょう。

jQueryのプレイグラウンドを作成します。これは後でjQueryのチャレンジで使います。

まず、テキストが`h3`の`jQuery Playground`要素を作成してください。

`h3`要素にブートストラップの`text-primary`クラスで色を付け、`text-center`クラスで中央揃えにしてください。

# --hints--

ページに`h3`要素を追加してください。

```js
assert.lengthOf(document.querySelectorAll('h3'),1);
```

`h3`要素には終了タグが必要です。

```js
assert.match(code,/<\/h3>/g);
assert.match(code,/<h3/g);
assert.equal( code.match(/<\/h3>/g).length , code.match(/<h3/g).length);
```

`h3`要素には`text-primary`クラスを適用して色を付けてください。

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-primary'));
```

`h3`要素には`text-center`クラスを適用して中央揃えにしてください。

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-center'));
```

`h3`要素にはテキスト`jQuery Playground`を入れてください。

```js
assert.match(document.querySelector('h3')?.textContent, /jquery(\s)+playground/gi);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```
