---
id: bad87fee1348bd9aec908746
title: ブートストラップの container-fluid div 内にページを配置する
challengeType: 0
forumTopicId: 18198
dashedName: house-our-page-within-a-bootstrap-container-fluid-div
---

# --description--

これでページ上のすべてのコンテンツがモバイル対応になるようにしましょう。

`h3` 要素を `div` クラスを持つ `container-fluid` 要素の中にネストしてください。

# --hints--

`div` 要素には `container-fluid` クラスを付けてください。

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container-fluid'));
```

`div` の各要素には終了タグが必要です。

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);

assert.equal(code.match(/<\/div>/g).length ,code.match(/<div/g).length);
```

`h3` 要素は `div` 要素の中にネストしてください。

```js
const divElement = document.querySelector('div');
const divChildren = divElement?.querySelectorAll(`:scope ${'h3'}`)
assert.lengthOf(divChildren,1);
```

# --seed--

## --seed-contents--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```

# --solutions--

```html
<div class="container-fluid">
    <h3 class="text-primary text-center">jQuery Playground</h3>
</div>
```
