---
id: 686386e76197458b7dfd52ef
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

现在你将创建一个 CTA 按钮。行动号召（CTA）是一个按钮或链接，旨在提示用户执行特定操作，在本例中是开始收听。

在最内层的 `div` 外，创建一个 `href` 为 `#` 且文本为 `Start listening` 的锚点元素。

这就是第一个定价卡的超文本标记语言代码。

# --hints--

你应该创建一个锚元素。

```js
assert.exists(document.querySelector("a"))
```

你的锚点元素应将 `href` 设置为 `#`。

```js
const link = document.querySelector("a")
assert.equal(link?.getAttribute("href"), "#")
```

你的锚点元素应包含文本 `Start listening`。

```js
assert.equal(document.querySelector("a")?.textContent, "Start listening")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          --fcc-editable-region--
          
          --fcc-editable-region-- 
        </div>
      </div>
    </main>
  </body>
</html>
```
