---
id: 686383abac6c625a84458ea9
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次研讨会中，你将通过使用 Tailwind CSS 的网格实用类来构建一个定价组件，从而深入了解 Tailwind CSS。

基本的超文本标记语言 公式化 已为你提供。首先创建一个带有文本 `Choose your listening plan` 的 `h1` 元素。

# --hints--

你应该创建一个 `h1` 元素。

```js
assert.exists(document.querySelector("h1"));
```

你的 `h1` 元素应包含文本 `Choose your listening plan`。

```js
assert.equal(document.querySelector("h1")?.textContent, "Choose your listening plan");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </main>
  </body>
</html>
```
