---
id: 686386e76197458b7dfd52eb
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

在 `h1` 下创建一个 `div` 元素。在该 `div` 内，创建另一个 `div` 元素，并在其内部嵌套另一个 `div` 元素。

并且在最内层的 `div` 内，创建一个带有文本 `Listener` 的 `h2` 元素。这将表现第一个定价方案。

# --hints--

你应该创建一个 `div` 元素。

```js
assert.exists(document.querySelector("div"));
```

你应该在外层的 `div` 内部创建另一个 `div` 元素。

```js
assert.exists(document.querySelector("div > div"));
```

你应该在内部的 `div` 元素内创建另一个 `div` 元素。

```js
assert.exists(document.querySelector("div > div > div"));
```

你应该在最内层的 `div` 元素内创建一个带有 `Listener` 文本的 `h2`。

```js
assert.equal(document.querySelector("div > div > div > h2")?.textContent, "Listener");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
