---
id: 686386e76197458b7dfd52ee
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

在 `ul` 内，创建两个列表项。在第一个列表项内，创建一个 `span` 元素并将超文本标记语言实体 `&#10003;` 作为其文本添加。紧接着该 `span` 但仍在列表项内，添加文本 `Ad-supported streaming`。

在第二个列表项内，创建一个 `span` 元素并将实体 `&#10003;` 作为其文本添加。在该 `span` 之后，但仍在列表项内，添加文本 `Curated playlists`。

# --hints--

你的 `ul` 元素内应有两个列表项。

```js
assert.lengthOf(document.querySelectorAll("ul > li"), 2)
```

你的第一个 `li` 元素应包含一个带有超文本标记语言字符实体 `&#10003;` 的 `span` 元素。

```js
const listItems = document.querySelectorAll("ul > li")
const firstSpan = listItems[0].querySelector("span")

assert.equal(firstSpan?.textContent, "✓")
```

你的第一个 `li` 元素应在 `span` 元素之后包含文本 `Ad-supported streaming`。

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[0].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Ad-supported streaming");
```

你的第二个 `li` 元素应包含一个带有超文本标记语言字符实体 `&#10003;` 的 `span` 元素。

```js
const listItems = document.querySelectorAll("ul > li")
const secondSpan = listItems[1].querySelector("span")

assert.equal(secondSpan?.textContent, "✓")
```

你的第二个 `li` 元素应在 `span` 元素之后包含文本 `Curated playlists`。

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[1].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Curated playlists");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>Start exploring millions of songs with basic features and ads.</p>
            <ul>
            --fcc-editable-region--
              
            --fcc-editable-region--
            </ul>
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
