---
id: 686386e76197458b7dfd52ec
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在 `h2` 下，创建一个带有文本 `$0` 的 `p` 元素。然后，在该 `p` 元素内且文本之后，创建一个带有文本 `/month` 的 `span` 元素。

# --hints--

你应该创建一个带有文本 `$0` 的 `p` 元素。

```js
const p = document.querySelector("p");
const pTextBeforeSpan = p.childNodes[0].textContent.trim();

assert.equal(pTextBeforeSpan, "$0");
```

你应该在 `p` 元素内创建一个包含文本 `/month` 的 `span` 元素。

```js
const spanText = document.querySelector("p > span")?.textContent;
assert.equal(spanText, "/month");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            --fcc-editable-region--
            
            --fcc-editable-region--
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
