---
id: 687d8e05a529a4838f1c961c
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

接著，移除 `Node` 類別內的 `pass` 關鍵字，然後為 `Node` 類別新增一個帶有兩個參數 `self` 和 `element` 的 `__init__` 方法。

在該方法中，將 `element` 指定給 `self.element`。這會儲存建立節點時傳入的值。

# --hints--

你應該在你的 `Node` 類別中定義一個 `__init__` 方法。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").has_function("__init__")`) 
})
```

你的 `__init__` 方法應該有兩個參數，`self` 和 `element`。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").find_function("__init__").has_args("self, element")`) 
})
```

你應該在 `__init__` 方法中將 `element` 指定到 `self.element`。

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.element == "A"
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
--fcc-editable-region--
    class Node:
        pass
--fcc-editable-region--
```
