---
id: 687d7deea6a2a22fa8816a73
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在之前的課程中，你學習了一些常見的資料結構。在這個工作坊中，你將實作一個鏈結串列。

首先，創建一個名為 `LinkedList` 的類別來表述將會保存你的列表節點的結構。在類別內添加 `pass` 關鍵字，以便測試能通過。

# --hints--

你應該創建一個名為 `LinkedList` 的類別。記得為你的程式碼添加 `pass` 關鍵字。

```js
({
    test: () => runPython(`assert _Node(_code).has_class("LinkedList")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
