---
id: 687d8c6f066877768f7ee441
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

在你的 `LinkedList` 類別中，將 `pass` 關鍵字取代為一個名為 `Node` 的內隱類別。這將表述儲存在列表中的個別節點資料。

# --hints--

你應該在 `LinkedList` 類別中創建一個 `Node` 類別。記得在 `Node` 類別內添加 `pass` 關鍵字。

```js
({
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").has_class("Node")`) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class LinkedList:
    pass
--fcc-editable-region--
```
