---
id: 687e366b3004a9831e4cbd41
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

為你的程式碼添加一個名為 `is_empty` 的方法，該方法傳回 `self.length == 0`。

你會用這個來追蹤列表中有多少節點。所以，當 `self.length == 0` 時，它會傳回 `True`，表示列表中沒有節點。 

如果 `self.length` 大於 `0`，它會傳回 `False`，表示該列表至少有一個節點。

# --hints--

你應該創建一個 `is_empty` 方法，該方法傳回 `self.length == 0`。

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.is_empty() == True
    ll.length = 1
    assert ll.is_empty() == False`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

--fcc-editable-region--
    
--fcc-editable-region--
```
