---
id: 687e344cc1993f7316af030c
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

在 `LinkedList` 類別中，為你的 `__init__` 方法添加一個將 `self.length` 設定為 `0` 的程式碼。這會初始化一個計數器，用來追蹤目前列表中有多少節點（元素）。

# --hints--

你應該在 `LinkedList` 類別的 `__init__` 方法中將 `self.length` 設定為 `0`。

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.length == 0`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    
--fcc-editable-region--
```
