---
id: 68800a64b06e80d9a31c131b
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

仍然在 `LinkedList` 類別的 `__init__` 方法中，將 `self.head` 設定為 `None`。這作為列表中第一個節點的引用。列表目前是空的，所以你必須將它設定為 `None`。


# --hints--

你應該在 `LinkedList` 類別的 `__init__` 方法中將 `self.head` 設為 `None`。

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.head is None`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    def __init__(self):
        self.length = 0
--fcc-editable-region--
```
