---
id: 688003f589b875af60273dd9
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

仍然在 `__init__` 方法內，將 `None` 指定給 `self.next`。這會將 `next` 指標設為 `None`，表示新節點尚未指向列表中的任何其他節點。

稍後，`next` 可以更新為指向另一個節點，以形成鏈結串列的串鏈。

# --hints--

在 `__init__` 方法中，你應該將 `next` 屬性初始化為 `None`。

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.next is None
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
--fcc-editable-region--
        def __init__(self, element):
            self.element = element
--fcc-editable-region--
```
