---
id: 68857e0a49baf4751f737a2a
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

現在，在 `is_empty` 方法之後，於 `LinkedList` 類別內建立一個名為 `add` 的方法。新的方法應該有兩個參數，`self` 和 `element`。

在你的新方法中，創建一個名為 `node` 的變數，並使用 `self.Node(element)` 指定一個新的 `Node` 實體。這將創建一個新的節點以添加到鏈結串列中。

# --hints--

你應該在 `LinkedList` 類別內創建一個 `add` 方法。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").has_function("add")`
      )
    )
})
```

你的 `add` 方法應該有兩個參數，`self` 和 `element`。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("add").has_args("self, element")`) 
})
```

你應該在 `add` 方法中將 `self.Node(element)` 指定給 `node`。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_body().is_equivalent("node = self.Node(element)")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
