---
id: 68867c3898b09115d8c5140f
title: 步驟 13
challengeType: 20
dashedName: step-13
---

# --description--

你剛剛創建的 `while` 迴圈讓你能夠遍歷列表中的每個節點，直到到達最後一個節點。當發生這種情況時，你可以將節點添加到列表。

在 `while` 區塊之後、`else` 區塊中，將 `node` 指定為 `current_node.next`。

# --hints--

你應該在 `else` 區段中將 `node` 指定給 `current_node.next`。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].has_stmt("current_node.next = node")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
