---
id: 688649c8673f842f49228a5e
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

如果列表不為空，該方法應該找到最後一個節點並將新節點連結到它的 `next` 指標。

為此，在 `if` 區塊之後創建一個 `else` 區塊。在 `else` 區塊內，創建一個 `current_node` 變數並將 `self.head` 指定給它。

# --hints--

你應該在 `else` 區段內將 `self.head` 指定給 `current_node`。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].is_equivalent("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
