---
id: 6888f9827637322d4f309192
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

在 `remove` 方法中，創建一個名為 `current_node` 的變數並將其指定為 `self.head` 的值。

這會從第一個節點開始遍歷鏈結串列。`current_node` 將用來追蹤你在尋找包含指定元素以移除的節點時，當前正在檢查的節點。

# --hints--

你應該在 `remove` 方法中將 `self.head` 指定給 `current_node`。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("remove").has_stmt("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

--fcc-editable-region--
    def remove(self, element):
        previous_node = None
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
