---
id: 6888d05ba074461b64f8eaca
title: 第 16 步
challengeType: 20
dashedName: step-16
---

# --description--

接下來，你將創建另一個方法，用於從鏈結串列中移除指定的元素。

創建一個名為 `remove` 的方法，帶有兩個參數 `self` 和 `element`。在其中，創建一個名為 `previous_node` 的變數並指定 `None`。

# --hints--

你應該創建一個名為 `remove` 的方法。

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_class("LinkedList").has_function("remove")
    `)
  )
})
```

你的 `remove` 方法應該有兩個參數，`self` 和 `element`。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("remove").has_args("self, element")`) 
})
```

你應該在 `remove` 方法中將 `None` 指定給 `previous_node`。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("remove").find_body().is_equivalent("previous_node = None")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
