---
id: 688acb2053ac7978a06345f1
title: 第 21 步
challengeType: 20
dashedName: step-21
---

# --description--

接著，創建一個 `elif` 述語來檢查 `previous_node` 是否不是 `None`。在其中將 `current_node.next` 指定為 `previous_node.next`。

如此一來，如果找到要移除的元素且它不在頭節點，前一節點的 `next` 指標將會更新，以跳過目前節點（要移除的元素），將其從鏈結串列中移除。

# --hints--

你應該創建一個 `elif` 述語來檢查 `previous_node` 是否不是 `None`。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("previous_node is not None")
`) })
```

你應該在 `elif` 述語中將 `current_node.next` 指定給 `previous_node.next`。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0]
assert _cond.is_equivalent("if current_node is None:\\n return\\nelif previous_node is not None:\\n previous_node.next = current_node.next")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
        while current_node is not None and current_node.element != element:
            previous_node = current_node
            current_node = current_node.next
--fcc-editable-region--
        if current_node is None:
            return
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
