---
id: 688ace81e3fb5a95a1f04c55
title: 第 22 步
challengeType: 20
dashedName: step-22
---

# --description--

接著，創建一個 `else` 區段，將 `current_node.next` 指定給 `self.head`。

這個控制代碼處理要移除的元素位於頭節點的情況。透過將 `self.head` 更新為指向 `current_node.next`，頭節點就會從鏈結串列中移除。

# --hints--

你應該創建一個 `else` 區塊，該區塊將 `current_node.next` 指定給 `self.head`。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0]
assert _cond.is_equivalent("if current_node is None:\\n return\\nelif previous_node is not None:\\n previous_node.next = current_node.next\\nelse:\\n self.head = current_node.next")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
        while current_node is not None and current_node.element != element:
            previous_node = current_node
            current_node = current_node.next
--fcc-editable-region--
        if current_node is None:
            return        
        elif previous_node is not None:
            previous_node.next = current_node.next
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
