---
id: 688ac806194cfc60ba1e1767
title: 第 20 步
challengeType: 20
dashedName: step-20
---

# --description--

在 `while` 迴圈之後，創建一個 `if` 述語來檢查 `current_node` 是否為 `None`。

當發生這種情況時，鏈結串列已被巡訪但未找到要移除的元素，函式應該停止執行。因此，請在 `if` 區塊內添加 `return` 關鍵字。

# --hints--

你應該創建一個 `if` 述語來檢查 `current_node` 是否為 `None`，並在其中添加 `return` 關鍵字。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0]
assert _cond.is_equivalent("if current_node is None:\\n return")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
        while current_node is not None and current_node.element != element:
            previous_node = current_node
            current_node = current_node.next
--fcc-editable-region--
        
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
