---
id: 68b84df2648b4c96b25022f3
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

現在是時候測試你目前的鏈結串列了！在你的程式碼底部，創建一個名為 `LinkedList` 的 `my_list` 實體。然後，使用你在上一步創建的 `is_empty` 方法，列印 `my_list.is_empty()`。

# --hints--

你應該創建一個名為 `LinkedList` 的 `my_list` 實體。

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_variable("my_list").value_is_call("LinkedList")
    `)
  )
})
```

你應該列印 `my_list.is_empty()`。

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).has_call("print(my_list.is_empty())")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0

--fcc-editable-region--

--fcc-editable-region--
```
