---
id: 68b9e951e4892721632acc26
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

太好了！現在讓我們測試為你的鏈結串列添加元素。在你的程式碼底部，在建立名為 `my_list` 的 `LinkedList` 實體並列印 `my_list.is_empty()` 之後，使用 `add` 方法將 `1` 和 `2` 添加到列表中。

然後，再次列印 `my_list.is_empty()` 以確認列表不再是空的。最後，列印 `my_list.length` 以查看鏈結串列的當前長度。

# --hints--

你應該使用 `add` 方法將 `1` 添加到列表中。

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(1)")`)
  )
})
```

你應該使用 `add` 方法將 `2` 添加到列表中。

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(2)")`)
  )
})
```

你應該列印 `my_list.is_empty()`。

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[1].is_equivalent("print(my_list.is_empty())")
    `)
  )
})
```

你應該列印 `my_list.length`。

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[2].is_equivalent("print(my_list.length)")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

my_list = LinkedList()
print(my_list.is_empty())

--fcc-editable-region--

--fcc-editable-region--
```
