---
id: 685239f0cad43c7a909cbd98
title: معالجة الأخطاء في صفحة تبني الحيوانات الأليفة
challengeType: 25
dashedName: lab-debug-pet-adoption-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

سالي، مالكة مخزن تبني الحيوانات، أنشأت أول صفحة ويب لها، لكن هناك بعض المشاكل.

مهمتك هي إصلاح كل الأخطاء حتى تتمكن سالي من متابعة بناء صفحتها.

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. تريد سالي استخدام صورة لبعض القطط لكنها لا تظهر بشكل صحيح. ستحتاج إلى إصلاح التالي في عنصر `img`:
   - استبدل الخاصية `href` بالخاصية الصحيحة لمصدر الصورة.
   - استبدل الخاصية `att` بالخاصية الصحيحة التي تمثل نصًا وصفيًا قصيرًا للصور.
   - أزل وسم الإغلاق `</img>` لأن عناصر `img` هي عناصر فارغة ولا تحتوي على وسم إغلاق.
2. تريد سالي استخدام بعض الروابط لتوجيه المستخدمين إلى صفحات الكلاب والقطط. لكن الروابط لا تعمل بشكل صحيح. ستحتاج إلى إصلاح التالي في عناصر `a`:
   - استبدل كلا الخاصيتين `src` بالخاصية الصحيحة المستخدمة لتحديد عناوين URL.

# --hints--

يجب أن يحتوي عنصر `img` على خاصية `src` بدلًا من الخاصية `href`.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("src"));
```

يجب أن يحتوي عنصر `img` على خاصية `alt` بدلًا من الخاصية غير الموجودة `att`.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("alt"));
```

يجب ألا يحتوي عنصر `img` على وسم الإغلاق `</img>`.

```js
assert.notMatch(code, /<\/img>/);
```

يجب أن يحتوي عنصر `a` الذي يحمل النص `Visit cats page` على خاصية `href` بدلًا من خاصية `src`.

```js
const anchors = document.querySelectorAll("a");
const catLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit cats page");
assert.isDefined(catLink);
assert.isTrue(catLink.hasAttribute("href"));
```

يجب أن يحتوي عنصر `a` الذي يحمل النص `Visit dogs page` على خاصية `href` بدلًا من خاصية `src`.

```js
const anchors = document.querySelectorAll("a");
const dogLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit dogs page");
assert.isDefined(dogLink);
assert.isTrue(dogLink.hasAttribute("href"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" att="Two tabby kittens sleeping together on a couch."></img>

<h2>Adopt a cat!</h2>
<a src="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a src="/dogs">Visit dogs page</a>
```

# --solutions--

```html
<h1>Welcome to XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">

<h2>Adopt a cat!</h2>
<a href="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a href="/dogs">Visit dogs page</a>
```
