---
id: 5cd9a70215d3c4e65518328f
title: بناء عد تنازلي
challengeType: 26
dashedName: build-a-countdown
---

# --description--

في هذا المختبر، ستبني دالة `countdown` تأخذ رقمًا ابتدائيًا وتُولد قائمة عد تنازلي.

عند استدعاء `countdown(n)`، يجب أن تُرجع مصفوفة تبدأ من `n` وتعد تنازليًا بمقدار `1` حتى تصل إلى `1`.

على سبيل المثال:

- يجب أن تُرجع `countdown(5)` `[5, 4, 3, 2, 1]`
- يجب أن تُرجع `countdown(1)` `[1]`

إذا كان الرقم الابتدائي أقل من `1`، فلا يوجد شيء للعد التنازلي، لذا أرجع مصفوفة فارغة:

- يجب أن تُرجع `countdown(0)` `[]`
- يجب أن تُرجع `countdown(-3)` `[]`

لإكمال المختبر، يجب أن تبني النتيجة باستخدام العودية ويجب ألا تستخدم الحلقات (`for`، `while`)، أو `Array.from()`، أو طرق تكرار المصفوفة (`forEach`، `map`، `filter`، `reduce`). يجب أن ينتج كل استدعاء مصفوفة نتيجته الخاصة (لا تستخدم متغيرات عامة لتخزين العد التنازلي).

**الهدف:** حقق قصص المستخدم أدناه واجتز جميع الاختبارات لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ دالة باسم `countdown`.
2. يجب أن تأخذ دالة `countdown` معلمة واحدة `n` (رقم).
3. إذا كان `n` أقل من `1`، يجب أن تُرجع `countdown` مصفوفة فارغة.
4. خلاف ذلك، يجب أن تُرجع `countdown` مصفوفة تحتوي على الأعداد الصحيحة من `n` نزولًا إلى `1` بترتيب تنازلي.
5. يجب أن تستخدم العودية (يجب أن تستدعي الدالة نفسها) ويجب ألا تستخدم أي نوع من الحلقات (`for`، `while`، `Array.from()`، أو طرق رتبة عليا مثل `forEach`، `map`، `filter`، `reduce`).
6. يجب أن يستخدم كل استدعاء عودي قيمة أصغر من الاستدعاء السابق (مثلًا، `n - 1`) حتى تصل الدالة إلى الحالة الأساسية.
7. يجب ألا تستخدم الحل متغيرات عامة لتخزين أو تخزين مؤقت لمصفوفة النتيجة.
8. يجب أن يُرجع استدعاء `countdown` عدة مرات مع مُدخلات مختلفة دائمًا النتيجة الصحيحة والمستقلة.

# --hints--

يجب أن يكون لديك دالة باسم `countdown`.

```js
assert.isFunction(countdown);
```

يجب أن تُرجع `countdown(-1)` مصفوفة فارغة.

```js
assert.isEmpty(countdown(-1));
```

يجب أن تُرجع `countdown(10)` `[10, 9, 8, 7, 6, 5, 4, 3, 2, 1]`

```js
assert.deepStrictEqual(countdown(10), [10, 9, 8, 7, 6, 5, 4, 3, 2, 1]);
```

يجب أن تُرجع `countdown(5)` `[5, 4, 3, 2, 1]`

```js
assert.deepStrictEqual(countdown(5), [5, 4, 3, 2, 1]);
```

يجب ألا يعتمد شفرتك على أي نوع من الحلقات (`for`، `while`، `Array.from()` أو طرق رتبة عليا مثل `forEach`، `map`، `filter`، و`reduce`).

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce|Array\.from/g)
);
```

يجب أن تستخدم العودية لحل هذه المشكلة.

```js
assert.isAtLeast(
  __helpers.removeJSComments(countdown.toString()).match(/countdown\s*\(.+\)/g)?.length,
  2
);
```

يجب ألا تُستخدم المتغيرات العامة لتخزين المصفوفة مؤقتًا.

```js
countdown(1)
assert.deepStrictEqual(countdown(5), [5, 4, 3, 2, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function countdown(n) {
  if (n < 1) {
    return [];
  }

  return [n, ...countdown(n - 1)];
}
```
