---
id: 669fdc11c9b0263fe0814a7a
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

각 `li` 요소 안에는 앵커 요소가 있어야 합니다.

첫 번째 앵커 요소의 텍스트는 `About`이어야 하며 `href` 속성 값은 `"#about"`이어야 합니다. `about` 앞의 해시 기호는 프로젝트에서 나중에 추가될 `id` 이름을 나타냅니다.

두 번째 앵커 요소의 텍스트는 `Posts`이어야 하며 `href` 속성 값은 `"#posts"`이어야 합니다.

세 번째 앵커 요소의 텍스트는 `Contact`이어야 하며 `href` 속성 값은 `"#contact"`이어야 합니다.

# --hints--

첫 번째 앵커 요소는 첫 번째 `About` 요소 안에 `li` 텍스트가 있어야 합니다.

```js
const anchorElement = document.querySelector("ul li:first-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "About");
```

첫 번째 앵커 요소는 `href` 속성이 `"#about"`로 설정되어야 합니다.

```js
const anchorElement = document.querySelector("ul li:first-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#about");
```

두 번째 앵커 요소는 두 번째 `Posts` 요소 안에 `li` 텍스트가 있어야 합니다.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
assert.strictEqual(anchorElement?.innerText.trim(), "Posts");
```

두 번째 앵커 요소는 `href` 속성이 `"#posts"`로 설정되어야 합니다.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#posts");
```

세 번째 앵커 요소는 세 번째 `Contact` 요소 안에 `li` 텍스트가 있어야 합니다.

```js
const anchorElement = document.querySelector("ul li:last-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "Contact");
```

세 번째 앵커 요소는 `href` 속성이 `"#contact"`로 설정되어야 합니다.

```js
const anchorElement = document.querySelector("ul li:last-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#contact");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
--fcc-editable-region--
          <li></li>
          <li></li>
          <li></li>
--fcc-editable-region--
        </ul>
      </nav>
    </header>
  </body>
</html>
```
