---
id: 669fde2081f65141ad703fe4
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

페이지의 첫 번째 섹션은 소개 섹션이 될 것입니다. 이 섹션은 Mr. Whiskers를 소개하고 이 블로그가 무엇에 관한 것인지 사용자에게 알려줍니다.

`main` 요소 안에 `section` 요소를 추가하고 `id` 속성을 `"about"`으로 설정하세요.

`section` 요소 안에 `h2`를 추가하고 `About` 텍스트를 넣으세요.

# --hints--

`section` 요소는 여는 태그가 있어야 합니다.

```js
assert.match(code, /<section.*>/);
```

`section` 요소는 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/section\>/);
```

`section` 요소는 `id` 속성이 `"about"`으로 설정되어야 합니다.

```js
const sectionElement = document.querySelector('section');
assert.strictEqual(sectionElement?.getAttribute("id"), "about");
```

`section` 요소는 `main` 요소 안에 있어야 합니다.

```js
const sectionElement = document.getElementById('about');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

`h2` 요소는 여는 태그가 있어야 합니다.

```js
assert.match(code, /<h2\>/);
```

`h2` 요소는 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/h2\>/);
```

`h2` 요소는 `About` 텍스트를 가져야 합니다. 철자를 다시 확인하세요.

```js
const h2Element = document.querySelector('body h2');
assert.strictEqual(h2Element?.innerText.trim(), "About");
```

`h2` 요소는 `section` 요소 안에 있어야 합니다.

```js
const h2Element = document.querySelector('#about h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
