---
id: 669fdfc9a5e3da42d2376609
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

about 섹션을 추가했으니, `About` 링크를 클릭해 페이지가 해당 섹션으로 이동하는지 확인하세요.

블로그 페이지의 다음 섹션은 Mr. Whiskers에 대해 이야기하는 게시물 목록이 될 것입니다.

`section` 요소를 하나 더 추가하고 `id`를 `"posts"`로 설정하세요.

`section` 요소 안에 `h2` 요소를 추가하고 `Posts` 텍스트를 넣으세요.

# --hints--

코드에 여는 `section` 태그가 있어야 합니다.

```js
assert.match(code, /<section.*>/i);
```

코드에 닫는 `section` 태그가 있어야 합니다.

```js
assert.match(code, /<\/section>/i);
```

`section` 요소에 `id`가 `"posts"`로 설정되어야 합니다.

```js
assert.match(code, /<section.*id="posts".*>/i);
```

`section` 요소는 `main` 요소 안에 있어야 합니다.

```js
const sectionElement = document.getElementById('posts');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

`h2` 요소는 `section` 요소 안에 있어야 합니다.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

`h2` 요소에는 `Posts` 텍스트가 있어야 합니다.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.innerText.trim(), 'Posts');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </main>
  </body>
</html>
```
