---
id: 669fd6fd12918e3de87854d4
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

블로그에서는 사용자가 페이지 내의 다른 섹션으로 이동할 수 있는 방법이 있어야 합니다.

`nav` 요소는 문서 내 다른 섹션이나 웹사이트 내 다른 섹션으로 이동하는 탐색 링크를 제공하는 데 사용됩니다. `nav` 요소가 메뉴나 목차에 자주 사용되는 것을 볼 수 있습니다.

다음은 `nav` 요소를 사용하는 예시입니다:

```html
<nav>
  <ul>
    <li><a href="#">Home</a></li>
    <li><a href="#about">About</a></li>
    <li><a href="#contact">Contact</a></li>
  </ul>
</nav>
```

`figure` 요소 아래에 `nav` 요소를 추가하고 그 안에 `ul` 요소를 중첩하세요.

`ul` 요소 안에 `li` 요소 세 개를 추가하세요.

# --hints--

`nav` 요소는 여는 태그가 있어야 합니다.

```js
assert.match(code, /<nav\>/);
```

`nav` 요소는 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/nav\>/);
```

`nav` 요소는 `header` 요소 안에 있어야 합니다.

```js
const navElement = document.querySelector('nav');
assert.strictEqual(navElement?.parentElement.tagName, 'HEADER');
```

`ul` 요소는 여는 태그가 있어야 합니다.

```js
assert.match(code, /<ul\>/);
```

`ul` 요소는 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/ul\>/);
```

`ul` 요소는 `nav` 요소 안에 있어야 합니다.

```js
const ulElement = document.querySelector('ul');
assert.strictEqual(ulElement?.parentElement.tagName, 'NAV');
```

`ul` 요소는 `li` 요소 세 개를 포함해야 합니다.

```js
const ulElement = document.querySelector("ul");

const liElements = ulElement?.querySelectorAll("li");
assert.strictEqual(liElements?.length, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
