---
id: 669fd003cb89ee3c2402e041
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

이 소개 내용에서는 Mr. Whiskers의 이미지와 캡션을 보여주세요.

`h1` 요소 아래에 `figure` 요소를 추가하는 것부터 시작하세요.

`figure` 요소 안에 `img` 요소를 추가하세요.

`src`의 `img` 속성은 `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"` 값을 가져야 하며 `alt` 텍스트는 `"a cat in the garden"` 값을 가져야 합니다.

`img` 요소 아래에 `figcaption`을 추가하고 텍스트는 `Mr. Whiskers in the Garden`로 하세요.

# --hints--

`figure` 요소는 여는 태그가 있어야 합니다. 

```js
assert.match(code, /<figure\>/);
```

`figure` 요소는 닫는 태그가 있어야 합니다. 

```js
assert.match(code, /<\/figure\>/);
```

`figure` 요소는 `header` 요소 안에 있어야 합니다.

```js
const figureElement = document.querySelector('figure');
assert.equal(figureElement?.parentElement.tagName, 'HEADER');
```

`img` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("img"));
```

`img` 요소는 `figure` 요소 안에 있어야 합니다.

```js
const imgElement = document.querySelector('img');
assert.strictEqual(imgElement?.parentElement.tagName, 'FIGURE');
```

`img` 요소는 `src` 속성이 `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`로 설정되어 있어야 합니다.

```js
const imgElement = document.querySelector('img');
const src = imgElement?.getAttribute('src');

assert.strictEqual(src, 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

`img` 요소는 `alt` 속성이 `"a cat in the garden"`로 설정되어 있어야 합니다.

```js
const imgElement = document.querySelector('img');
const alt = imgElement?.getAttribute('alt');

assert.strictEqual(alt, 'a cat in the garden');
```

`figcaption` 요소는 여는 태그가 있어야 합니다. 

```js
assert.match(code, /<figcaption\>/);
```

`figcaption` 요소는 닫는 태그가 있어야 합니다. 

```js
assert.match(code, /<\/figcaption\>/);
```

`figcaption` 요소는 `figure` 요소 안에 있어야 합니다.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.parentElement.tagName, 'FIGURE');
```

`figcaption` 요소는 `Mr. Whiskers in the Garden` 텍스트를 가져야 합니다.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.innerText.trim(), "Mr. Whiskers in the Garden")
```

`figcaption`은 `img` 요소 아래에 있어야 합니다. 순서가 잘못되었습니다.

```js
const figureElement = document.querySelector('figure');
assert.strictEqual(figureElement?.lastElementChild?.tagName, "FIGCAPTION")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
