---
id: 669fdd3965f36f40af9615e5
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

페이지 헤더 구성을 마쳤으니, 이제 본문 내용을 추가하기 시작하세요.

`header` 아래에 `main` 요소를 추가하세요.

# --hints--

`main` 요소는 여는 태그가 있어야 합니다.

```js
assert.match(code, /<main\>/);
```

`main` 요소는 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/main\>/);
```

`main` 요소는 `header` 요소 아래에 있어야 합니다.

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
