---
id: 669fcb5106eeb13ab7afcd09
title: 1단계
challengeType: 0
dashedName: step-1
---

# --description--

이 워크숍에서는 Mr. Whiskers라는 고양이를 주제로 한 블로그 페이지를 만들어 의미론적 HTML을 연습합니다.

먼저 만들 섹션은 페이지의 <dfn>헤더</dfn>입니다.

`header` 요소는 페이지 탐색과 기타 소개 정보를 나타내는 데 사용됩니다.

다음은 `header` 요소를 사용한 예시입니다:

```html
<header>
  <h1>Main Page Title Goes Here</h1>
  <img src="example-logo.png" alt="Example logo" />
</header>
```

`body` 요소 안에 `header` 요소를 추가하세요.

# --hints--

열린 `header` 태그가 있어야 합니다.

```js
assert.match(code, /<header>/i);
```

닫힌 `header` 태그가 있어야 합니다.

```js
assert.match(code, /<\/header>/i);
```

`header` 요소는 `body` 요소 안에 있어야 합니다.

```js
assert.exists(document.querySelector('body > header'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </body>
</html>
```
