---
id: 66a338be7f1dee383a0e0ecb
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

두 번째 블로그 게시물에는 또 다른 `article` 요소를 추가해야 합니다.

`article` 요소 안에 `h3` 요소를 추가하고 `Mr. Whiskers' First Bath` 텍스트를 넣으세요.

`h3` 요소 아래에 lorem ipsum 텍스트 두 단락을 추가하세요.

# --hints--

`article` 요소 안에 두 번째 `section` 요소가 있고 `id`이 `posts`로 설정되어 있어야 합니다.

```js
const articleElement = document.querySelector('#posts article:nth-of-type(2)');
assert.strictEqual(articleElement?.parentElement.tagName, 'SECTION');
```

`h3` 요소 안에 `article` 요소가 있고 `Mr. Whiskers' First Bath` 텍스트가 있어야 합니다.

```js
const h3Element = document.querySelector('#posts article:nth-of-type(2) h3');
assert.strictEqual(h3Element?.innerText.trim(), "Mr. Whiskers' First Bath");
```

`h3` 요소는 `article` 요소 안에 있어야 합니다.

```js
const h3Element = document.querySelector('#posts article:nth-of-type(2) h3');
assert.strictEqual(h3Element?.parentElement.tagName, 'ARTICLE');
```

`article` 요소 안에 두 개의 단락 요소가 있어야 합니다.

```js
const paragraphElements = document.querySelectorAll('#posts article:nth-of-type(2) p');
assert.strictEqual(paragraphElements.length, 2);
```

첫 번째 단락 요소에는 lorem ipsum 텍스트가 있어야 합니다. 맞춤법과 구두점 오류를 다시 확인하세요.

```js
const paragraphElement = document.querySelector('#posts article:nth-of-type(2) p:first-of-type');
const loremIpsum = `Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.`

assert.strictEqual(paragraphElement?.innerText.trim(), loremIpsum);
```

두 번째 단락 요소에도 lorem ipsum 텍스트가 있어야 합니다. 맞춤법과 구두점 오류를 다시 확인하세요.

```js
const paragraphElement = document.querySelector('#posts article:nth-of-type(2) p:last-of-type');
const loremIpsum = `Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.`

assert.strictEqual(paragraphElement?.innerText.trim(), loremIpsum);
```

단락들은 `h3` 요소 뒤에 와야 합니다.

```js
let h3 = document.querySelectorAll('#posts article:nth-of-type(2) > p ~ h3');
assert.isEmpty(h3);
const paragraphs = document.querySelectorAll('#posts article:nth-of-type(2) > h3 ~ p');
assert.lengthOf(paragraphs, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>

        <article>
          <h3>Mr. Whiskers' First Day Home</h3>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
        </article>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
