---
id: 66a334a1a7cca6354999f9bf
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

첫 번째 블로그 게시물에는 <dfn>article</dfn> 요소를 사용하세요.

`article` 요소는 웹 페이지에서 독립적인 내용을 나타냅니다.

```html
<article>
  <h1>Example heading</h1>
  <p>Example article text</p>
</article>
```

`h2` 요소 아래에 `article` 요소를 추가하세요.

`article` 요소 안에 텍스트 `h3`가 있는 `Mr. Whiskers' First Day Home` 요소를 추가하세요.

여기서 `h3`를 사용하는 이유는 제목 요소의 올바른 구조적 계층을 유지하는 것이 중요하기 때문입니다. 게시물의 부제목이 `h2` 요소이므로, 계층에서 다음 단계는 `h3`가 됩니다.

# --hints--

코드에 `article` 여는 태그가 있어야 합니다.

```js
assert.match(code, /<article>/i);
```

코드에 `article` 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/article>/i);
```

`article` 요소는 `section` 요소 안에 있어야 하며 `id`은 `posts`로 설정되어야 합니다.

```js
const articleElement = document.querySelector('#posts article');
assert.strictEqual(articleElement?.parentElement.tagName, 'SECTION');
```

코드에 `h3` 여는 태그가 있어야 합니다.

```js
assert.match(code, /<h3>/i);
```

코드에 `h3` 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/h3>/i);
```

`h3` 요소는 `Mr. Whiskers' First Day Home` 텍스트를 포함해야 합니다.

```js
const h3Element = document.querySelector('h3');
assert.strictEqual(h3Element?.innerText.trim(), "Mr. Whiskers' First Day Home");
```

`h3` 요소는 `article` 요소 안에 있어야 합니다.

```js
const h3Element = document.querySelector('#posts article h3');
assert.strictEqual(h3Element?.parentElement.tagName, 'ARTICLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
