---
id: 66a49f685961e997e337cab1
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`header`는 블로그의 주요 제목, 이미지, 페이지 내비게이션을 표시하는 역할을 합니다.

`header` 요소 안에 `h1`를 추가하고 `Welcome to Mr. Whiskers' Blog Page!` 텍스트를 넣으세요.

# --hints--

열린 `h1` 태그가 있어야 합니다.

```js
assert.match(code, /<h1>/i);
```

닫힌 `h1` 태그가 있어야 합니다.

```js
assert.match(code, /<\/h1>/i);
```

`h1` 요소에 `Welcome to Mr. Whiskers' Blog Page!` 텍스트가 있어야 합니다. 철자 오류가 없는지 다시 확인하세요.

```js
const h1Text = document.querySelector("h1")?.innerText.trim();
assert.strictEqual(h1Text, "Welcome to Mr. Whiskers' Blog Page!");
```

`h1` 요소는 `header` 요소 안에 중첩되어야 합니다.

```js
const h1Element = document.querySelector("h1");
assert.strictEqual(h1Element?.parentElement.tagName, 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
