---
id: 61fd5a93fd62bb35968adeab
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Configura tu HTML con los elementos `DOCTYPE`, `html` indicando que este documento está en inglés, `head` y `body`.

Dale a tu elemento `head` los elementos `meta` apropiados para el `charset` y el `viewport`, un elemento `title` con un titulo apropiado y un elemento `link` para tu hoja de estilos.

# --hints--

Tu código debe tener una declaración `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

El código debe tener un elemento `html`.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

Tu elemento `html` debe tener una etiqueta de apertura con el atributo `lang` con el valor `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Su elemento `html` debe tener una etiqueta de cierre.

```js
assert(code.match(/<\/html\s*>/));
```

Tu código debe tener un elemento `head` dentro del elemento `html`.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

Tu código debe tener un elemento `body` dentro del elemento `html`.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

Tu elemento `head` debe estar antes de tu elemento `body`.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

Debes tener dos elementos `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

Un elemento `meta` debe tener un `name` establecido en `viewport`, y el `content` establecido a `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

El otro elemento `meta` debe tener el atributo `charset` establecido en `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Tu código debe tener un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Tu `title` debe tener un poco de texto.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

Su código debe tener un elemento <code>link`.

```js
assert(/<link/.test(code))
```

Su elemento `link` debe estar dentro de su elemento `head`.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i))
```

Su elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

Tu elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
