---
id: 668f08ea07b99b1f4a91acab
title: 레시피 페이지 만들기
challengeType: 25
dashedName: build-a-recipe-page
demoType: onClick
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `!DOCTYPE html` 선언이 있어야 합니다.
1. `html`가 `lang`로 설정된 `en` 요소가 있어야 합니다.
1. 레시피 이름이 들어간 `head` 요소와 `title` 속성이 `meta`로 설정된 `charset` 요소를 포함하는 `UTF-8` 요소가 있어야 합니다.
1. `body` 요소가 있어야 합니다.
1. 레시피 이름이 들어간 `h1` 요소가 있어야 합니다.
1. `p` 아래에 레시피를 소개하는 `h1` 요소가 있어야 합니다.
1. 재료 섹션을 위한 텍스트 `h2`가 있는 `Ingredients` 요소가 하나 있어야 합니다.
1. 첫 번째 `ul` 요소 아래에 재료를 나열하는 최소 네 개의 항목(`li` 요소)을 포함한 순서 없는 목록(`h2` 요소)이 있어야 합니다.
1. 지침 섹션을 위한 텍스트 `h2`가 있는 두 번째 `Instructions` 요소가 있어야 합니다.
1. 두 번째 `ol` 아래에 레시피 단계를 순서대로 나열하는 최소 네 개의 항목이 있는 순서 있는 목록(`h2` 요소)이 있어야 합니다.
1. 유효한 이미지가 설정된 `img` 속성을 가진 `src` 요소가 하나 있어야 하며, 원한다면 `https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg`를 사용할 수 있고, 이미지를 설명하는 `alt` 속성도 있어야 합니다.

# --hints--

레시피 페이지에는 `!DOCTYPE html` 선언이 있어야 합니다.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

`html`가 `lang`로 설정된 `en` 요소가 있어야 합니다.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/gi);
```

`head` 요소 안에 `html` 요소가 있어야 합니다.

```js
assert.match(code, /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i);
```

`title` 요소 안에 `head` 요소가 있어야 합니다.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i);
```

`title` 요소에는 레시피 제목이 있어야 합니다.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

`meta` 요소 안에 `head` 요소가 있어야 합니다.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*meta[\s\S]*>[\s\S]*<\/\s*head\s*>/i);
```

`meta` 요소의 `charset` 속성은 `UTF-8`로 설정되어야 합니다.

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

`body` 요소 안에 `html` 요소가 있어야 합니다.

```js
assert.match(code, /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i);
```

레시피 이름이 들어간 `h1` 요소가 있어야 합니다.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

`h1` 요소는 하나만 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

`p` 요소 아래에 `h1` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('h1')?.nextElementSibling, document.querySelector('p'));
```

첫 번째 `p` 요소는 레시피를 설명해야 합니다.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

첫 번째 `h2` 요소에는 텍스트 `Ingredients`가 있어야 합니다.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Ingredients');
```

첫 번째 `h2` 요소 아래에 순서 없는 목록 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('ul')?.previousElementSibling.tagName, 'H2');
```

순서 없는 목록에는 재료가 들어간 최소 네 개의 항목 요소가 있어야 합니다.

```js
const els = document.querySelectorAll('ul > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

두 번째 `h2` 요소에는 텍스트 `Instructions`가 있어야 합니다.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Instructions');
```

두 번째 `h2` 요소 아래에 순서 있는 목록 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('h2')?.[1]?.nextElementSibling?.tagName, "OL");
```

순서 있는 목록에는 지침이 들어간 최소 네 개의 항목 요소가 있어야 합니다.

```js
const els = document.querySelectorAll('ol > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

최소 하나의 `img` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('img'));
```

모든 `img` 요소는 유효한 `src` 속성과 값을 가져야 합니다.

```js
const img = document.querySelector('img');
const rawSrc = img?.getAttribute('src');
const resolvedSrc = img?.src;
const re = new RegExp(window.location.href, "ig");

assert.isAbove(rawSrc?.trim().length, 0, "The 'src' attribute must be explicitly set.");
assert.notMatch(resolvedSrc, re, "The 'src' should not start with the current page URL");

img.onload = () => {
  console.log('Image loaded successfully.');
};

img.onerror = (error) => {
  console.error('Image failed to load:', error);
  assert.fail("Your image's URL should be valid."); // Make the test instafail
};

if (img.complete) {
  img.onload && img.onload();
};
```

모든 `img` 요소는 이미지를 설명하는 `alt` 속성을 가져야 합니다.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Chocolate chip cookies recipe</title>
</head>

<body>
  <h1>Chocolate Chip Cookies</h1>
  <p>Welcome to the ultimate guide for making mini chocolate chip cookies! These bite-sized treats are perfect for
    satisfying your sweet tooth without overindulging. Follow this simple recipe to create delicious,
    crispy-on-the-outside, chewy-on-the-inside mini chocolate chip cookies that everyone will love.</p>
  <img src="https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg" alt="Ingredients for baking: three eggs, a bowl of flour, a glass of milk, and a whisk arranged on a wooden table.">
  <h2>Ingredients</h2>
  <ul>
    <li>1 cup all-purpose flour</li>
    <li>1/2 teaspoon baking soda</li>
    <li>1/4 cup unsalted butter, softened</li>
    <li>1/4 cup granulated sugar</li>
    <li>1/2 teaspoon vanilla extract</li>
    <li>1/2 cup mini chocolate chips</li>
  </ul>
  <h2>Instructions</h2>
  <ol>
    <li>Preheat your oven to 350°F (175°C) and line a baking sheet with parchment paper.</li>
    <li>In a bowl, whisk together the flour and baking soda.</li>
    <li>In another bowl, beat the butter, sugar, and vanilla extract until creamy.</li>
    <li>Gradually add the dry ingredients to the wet mixture, then fold in the mini chocolate chips.</li>
    <li>Drop small spoonfuls of dough onto the baking sheet.</li>
    <li>Bake for 8-10 minutes, then let cool before enjoying!</li>
  </ol>
</body>

</html>
```
