---
id: 686d749d6f32f708fdf2b7ad
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In einer vorherigen Lektion haben Sie etwas über `svg`-Elemente gelernt und wie sie in Projekten oft als Symbole verwendet werden. In realen Codebasen würden Sie normalerweise auf Icon-Bibliotheken zurückgreifen, sodass Sie `svg`-Elemente nicht von Grund auf neu erstellen müssen. In diesem Workshop bauen Sie jedoch ein Herzsymbol, um die wichtigsten Attribute kennenzulernen, die innerhalb eines `svg`-Elements verwendet werden.

Beginnen Sie damit, ein `svg`-Element auf der Seite zu erstellen.

# --hints--

Sie sollten ein `svg`-Element auf der Seite haben.

```js
const svg = document.querySelector('svg');
assert.exists(svg);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region-- 
    
    --fcc-editable-region--
  </body>
</html>
```
