---
id: 686daa7ed79ceacd0b264e80
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Sie kommen dem Ziel näher, schauen Sie sich nun dieses Beispiel an:

```html
<svg viewBox="0 0 50 50">
</svg>
```

Das `viewBox`-Attribut steuert, welcher Teil des Bildes innerhalb des SVG sichtbar ist.

- Die ersten beiden Zahlen (`0 0`) legen die Startposition des `viewBox` fest – die obere linke Ecke (x und y). 
- Die nächsten beiden Zahlen (`50 50`) definieren die Breite und Höhe des `viewBox`.

Setzen Sie das `viewBox`-Attribut auf `0 0 24 24`.

# --hints--

Sie sollten ein `viewBox`-Attribut haben.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
assert.exists(viewBox);
```

Sie sollten die x-Position von `viewBox` auf 0 setzen.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const x = viewBox?.trim().split(' ')[0];
assert.strictEqual(x, '0');
```

Sie sollten die y-Position von `viewBox` auf 0 setzen.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const y = viewBox?.trim().split(' ')[1];
assert.strictEqual(y, '0');
```

Sie sollten die Breite von `viewBox` auf 24 setzen.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const width = viewBox?.trim().split(' ')[2];
assert.strictEqual(width, '24');
```

Sie sollten die Höhe von `viewBox` auf 24 setzen.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const height = viewBox?.trim().split(' ')[3];
assert.strictEqual(height, '24');
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region--
    <svg width="24" height="24">
    --fcc-editable-region--
      <path
        d="M12 21s-6-4.35-9.33-8.22C-.5 7.39 3.24 1 8.4 4.28 10.08 5.32 12 7.5 12 7.5s1.92-2.18 3.6-3.22C20.76 1 24.5 7.39 21.33 12.78 18 16.65 12 21 12 21z"
      ></path>
    </svg>
  </body>
</html>
```
