---
id: 6143869bb45bd85e3b1926aa
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

在你的 `.heading` 元素中，創建一個 `header` 元素，並將 `class` 設定為 `hero`。

在該元素中，創建一個 `img` 元素，並將 `src` 設定為 `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`，`alt` 設定為 `freecodecamp logo`，以及 `class` 設定為 `hero-img`。

`img` 元素上的 `loading` 屬性可以設定為 `lazy`，以告訴瀏覽器直到需要時（例如，當使用者將影像捲動到視窗中）才去擷取影像資源。作為額外的利益，延遲載入的元素會在非延遲載入的元素載入完成後才開始載入－這表示網際網路連線較慢的使用者可以在不必等待影像載入的情況下檢視你的頁面內容。

為你的新 `img` 元素添加一個設定為 `lazy` 的 `loading` 屬性。

# --hints--

你應該創建一個 `header` 元素。

```js
assert.exists(document.querySelector('header'));
```

你的 `header` 元素應該位於你的 `.heading` 元素內。

```js
assert.equal(document.querySelector('header')?.parentElement?.className, 'heading');
```

你的 `header` 元素應該將 `class` 設定為 `hero`。

```js
assert.equal(document.querySelector('header')?.className, 'hero');
```

你應該創建一個 `img` 元素。

```js
assert.exists(document.querySelector('img'));
```

你的 `img` 元素應該位於你的 `header` 元素內。

```js
assert.equal(document.querySelector('img')?.parentElement?.localName, 'header');
```

你的 `img` 元素應該將 `src` 設定為 `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`。

```js
assert.equal(document.querySelector('img')?.getAttribute('src'), 'https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png');
```

你的 `img` 元素應該將 `alt` 設定為 `freecodecamp logo`。

```js
assert.equal(document.querySelector('img')?.getAttribute('alt'), 'freecodecamp logo');
```

你的 `img` 元素應該將 `loading` 屬性設為 `lazy`。

```js
assert.equal(document.querySelector('img')?.getAttribute('loading'), 'lazy');
```

你的 `img` 元素應該將 `class` 設定為 `hero-img`。

```js
assert.equal(document.querySelector('img')?.className, 'hero-img');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
