---
id: 61438b5b66d76a6264430f2a
title: 步驟 9
challengeType: 0
dashedName: step-9
---

# --description--

在每個新的 `a` 元素中，增加一個 `i` 元素，併爲它們提供以下類：

- 你的第一個 `i` 元素應該有 `fab fa-facebook-f` 類
- 你的第二個 `i` 元素應該有 `fab fa-twitter` 類
- 你的第三個 `i` 元素應該有 `fab fa-instagram` 類
- 你的第四個 `i` 元素應該有 `fab fa-linkedin-in` 類
- 你的第五個 `i` 元素應該有 `fab fa-youtube` 類

# --hints--

你應該有五個 `i` 元素。

```js
assert.lengthOf(document.querySelectorAll('i'), 5);
```

每個 `a` 元素應該只有一個 `i` 元素。

```js
const aaaaa = document.querySelectorAll('.social-icons a');
assert.lengthOf(aaaaa, 5);
aaaaa?.forEach(
  a => {
    assert.lengthOf(a?.children, 1);
    assert.equal(a?.children?.[0]?.localName, 'i');
  }
);
```

每個 `i` 元素都應該有一個包括 `fab` 的 `class` 屬性。

```js
const iiiii = document.querySelectorAll('i');
assert.lengthOf(iiiii, 5);
iiiii?.forEach(
  i => {
    assert.isTrue(i?.classList?.contains('fab'));
  }
);
```

第一個 `i` 元素應該有一個 `class` 屬性，其中包括 `fa-facebook-f`。

```js
assert.isTrue(document.querySelectorAll('i')?.[0]?.classList?.contains('fa-facebook-f'));
```

第二個 `i` 元素應該有一個 `class` 屬性，其中包括 `fa-twitter`。

```js
assert.isTrue(document.querySelectorAll('i')?.[1]?.classList?.contains('fa-twitter'));
```

第三個 `i` 元素應該有一個 `class` 屬性，其中包括 `fa-instagram`。

```js
assert.isTrue(document.querySelectorAll('i')?.[2]?.classList?.contains('fa-instagram'));
```

第四個 `i` 元素應該有一個`class` 屬性，其中包括 `fa-linkedin-in`。

```js
assert.isTrue(document.querySelectorAll('i')?.[3]?.classList?.contains('fa-linkedin-in'));
```

第五個 `i` 元素應該有一個 `class` 屬性，其中包括 `fa-youtube`。

```js
assert.isTrue(document.querySelectorAll('i')?.[4]?.classList?.contains('fa-youtube'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--
        <div class="social-icons">
          <a href="https://www.facebook.com/freecodecamp/">
          </a>
          <a href="https://twitter.com/freecodecamp/">
          </a>
          <a href="https://instagram.com/freecodecamp">
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
          </a>
          <a href="https://www.youtube.com/freecodecamp">
          </a>
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
