---
id: 614387cbefeeba5f3654a291
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

在你的 `header` 元素之後，創建一個 `div`，並將 `class` 設定為 `author`。

在該 `div` 中，創建一個 `p` 元素，並將 `class` 設為 `author-name`，給它設定文字 `By freeCodeCamp`。將 `freeCodeCamp` 部分包裹在一個 `a` 元素中，並將 `href` 設為 `https://freecodecamp.org`，`target` 設為 `_blank`。


在其下方，為你的第二個 `p` 元素添加類別 `publish-date` 和文字 `March 7, 2019`。

# --hints--

你應該創建一個新的 `div` 元素。

```js
assert.exists(document.querySelector('div'));
```

你的 `div` 元素應該位於你的 `header` 元素之後。

```js
assert.equal(document.querySelector('div')?.previousElementSibling?.localName, 'header');
```

你的 `div` 元素應該將 `class` 設定為 `author`。

```js
assert.equal(document.querySelector('div')?.className, 'author');
```

你應該創建兩個新的 `p` 元素。

```js
assert.lengthOf(document.querySelectorAll('p'), 3);
```

你的兩個新的 `p` 元素應該在你的 `div` 元素內。

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

你的第一個新的 `p` 元素應該設定 `class` 為 `author-name`。

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.className, 'author-name');
```

你的第一個新的 `p` 元素應該有文字 `By freeCodeCamp`。

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.innerText, 'By freeCodeCamp');
```

你的第二個新的 `p` 元素應該設定 `class` 為 `publish-date`。

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className, 'publish-date');
```

你的第二個新的 `p` 元素應該有文字 `March 7, 2019`。

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText, 'March 7, 2019');
```

你應該創建一個新的 `a` 元素。

```js
assert.exists(document.querySelector('a'));
```

你的 `a` 元素應該位於你的第一個新的 `p` 元素內。

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName, 'a');
```

你的 `a` 元素應該將 `href` 設定為 `https://freecodecamp.org`。

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href'), 'https://freecodecamp.org');
```

你的 `a` 元素應該將 `target` 設定為 `_blank`。

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target'), '_blank');
```

你的 `a` 元素應該包圍文字 `freeCodeCamp`。

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.innerText, 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
