---
id: 614385513d91ae5c251c2052
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在這個工作坊，你將使用 freeCodeCamp 課程的前一版本的新聞稿創建一個雜誌頁面。

開始添加一個 `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` 字體樣式表 `link` 元素，一個 `https://use.fontawesome.com/releases/v5.8.2/css/all.css` FontAwesome 樣式表 `link` 元素，還有一個 `./styles.css` 樣式表 `link` 元素。

你的字體樣式表將加載三種不同的字體：`Anton`、`Baskervville` 和 `Raleway`。

# --hints--

你的代碼應該包含三個 `link` 元素。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

你的 `link` 元素應該在你的 `head` 元素中。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

你的三個 `link` 元素都應該有一個值爲 `stylesheet` 的 `rel` 屬性。

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

有一個鏈接元素應將 `href` 設置爲 `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`。

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

有一個鏈接元素應將 `href` 設置爲 `https://use.fontawesome.com/releases/v5.8.2/css/all.css`。

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

有一個 `link` 元素應將 `href` 應設置爲 `styles.css`。

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
