---
id: 614e0e588f0e8a772a8a81a6
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

你的影像目前佔用很多空間。為了更清楚地看到你正在處理的內容，為 `img` 元素添加一個 `width` 屬性，值為 `400`。

當你完成 CSS 的工作後，你將會移除此項。

# --hints--

你的 `img` 元素應該有一個設定為 `400` 的 `width` 屬性。

```js
assert.equal(document.querySelector('img')?.getAttribute('width'), '400');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
