---
id: 56533eb9ac21ba0edf2244e2
title: 시저 암호
challengeType: 5
saveSubmissionToDB: true
forumTopicId: 16003
dashedName: caesars-cipher
---

# --description--

가장 간단하고 널리 알려진 <dfn>암호</dfn> 중 하나는 <dfn>시저 암호</dfn>로, <dfn>시프트 암호</dfn>라고도 합니다. 시프트 암호는 문자들의 의미가 일정한 양만큼 이동하는 방식입니다.

현대에서 흔히 쓰이는 예는 <a href="https://www.freecodecamp.org/news/how-to-code-the-caesar-cipher-an-introduction-to-basic-encryption-3bf77b4e19f7/" target="_blank" rel="noopener noreferrer nofollow">ROT13</a> 암호로, 문자의 값이 13자리만큼 이동합니다. 따라서 `A ↔ N`, `B ↔ O` 등이 됩니다.

ROT13으로 인코딩된 문자열을 입력받아 디코딩된 문자열을 반환하는 함수를 작성하세요.

모든 문자는 대문자입니다. 알파벳이 아닌 문자(예: 공백, 구두점)는 변환하지 말고 그대로 전달하세요.

# --hints--

`rot13("SERR PBQR PNZC")`는 문자열 `FREE CODE CAMP`로 디코딩되어야 합니다.

```js
assert(rot13('SERR PBQR PNZC') === 'FREE CODE CAMP');
```

`rot13("SERR CVMMN!")`는 문자열 `FREE PIZZA!`로 디코딩되어야 합니다.

```js
assert(rot13('SERR CVMMN!') === 'FREE PIZZA!');
```

`rot13("SERR YBIR?")`는 문자열 `FREE LOVE?`로 디코딩되어야 합니다.

```js
assert(rot13('SERR YBIR?') === 'FREE LOVE?');
```

`rot13("GUR DHVPX OEBJA SBK WHZCF BIRE GUR YNML QBT.")`는 문자열 `THE QUICK BROWN FOX JUMPS OVER THE LAZY DOG.`로 디코딩되어야 합니다.

```js
assert(
  rot13('GUR DHVPX OEBJA SBK WHZCF BIRE GUR YNML QBT.') ===
    'THE QUICK BROWN FOX JUMPS OVER THE LAZY DOG.'
);
```

# --seed--

## --seed-contents--

```js
function rot13(str) {
  return str;
}

rot13("SERR PBQR PNZC");
```

# --solutions--

```js
var lookup = {
  'A': 'N','B': 'O','C': 'P','D': 'Q',
  'E': 'R','F': 'S','G': 'T','H': 'U',
  'I': 'V','J': 'W','K': 'X','L': 'Y',
  'M': 'Z','N': 'A','O': 'B','P': 'C',
  'Q': 'D','R': 'E','S': 'F','T': 'G',
  'U': 'H','V': 'I','W': 'J','X': 'K',
  'Y': 'L','Z': 'M'
};

function rot13(encodedStr) {
  var codeArr = encodedStr.split("");  // String to Array
  var decodedArr = []; // Your Result goes here
  // Only change code below this line

  decodedArr = codeArr.map(function(letter) {
    if(lookup.hasOwnProperty(letter)) {
      letter = lookup[letter];
    }
    return letter;
  });

  // Only change code above this line
  return decodedArr.join(""); // Array to String
}
```
