---
id: a7f4d8f2483413a6ce226cac
title: 로마 숫자 변환기
challengeType: 5
saveSubmissionToDB: true
forumTopicId: 16044
dashedName: roman-numeral-converter
---

# --description--

주어진 숫자를 로마 숫자로 변환하세요.

| 로마 숫자       | 아라비아 숫자    |
|----------------|-----------------|
| M              | 1000            |
| CM             | 900             |
| D              | 500             |
| CD             | 400             |
| C              | 100             |
| XC             | 90              |
| L              | 50              |
| XL             | 40              |
| X              | 10              |
| IX             | 9               |
| V              | 5               |
| IV             | 4               |
| I              | 1               |

모든 로마 숫자 답변은 대문자로 작성하세요.

# --hints--

`convertToRoman(2)` 는 문자열 `II` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(2), 'II');
```

`convertToRoman(3)` 는 문자열 `III` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(3), 'III');
```

`convertToRoman(4)` 는 문자열 `IV` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(4), 'IV');
```

`convertToRoman(5)` 는 문자열 `V` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(5), 'V');
```

`convertToRoman(9)` 는 문자열 `IX` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(9), 'IX');
```

`convertToRoman(12)` 는 문자열 `XII` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(12), 'XII');
```

`convertToRoman(16)` 는 문자열 `XVI` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(16), 'XVI');
```

`convertToRoman(29)` 는 문자열 `XXIX` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(29), 'XXIX');
```

`convertToRoman(44)` 는 문자열 `XLIV` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(44), 'XLIV');
```

`convertToRoman(45)` 는 문자열 `XLV` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(45), 'XLV');
```

`convertToRoman(68)` 는 문자열 `LXVIII` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(68), 'LXVIII');
```

`convertToRoman(83)` 는 문자열 `LXXXIII` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(83), 'LXXXIII');
```

`convertToRoman(97)` 는 문자열 `XCVII` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(97), 'XCVII');
```

`convertToRoman(99)` 는 문자열 `XCIX` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(99), 'XCIX');
```

`convertToRoman(400)` 는 문자열 `CD` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(400), 'CD');
```

`convertToRoman(500)` 는 문자열 `D` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(500), 'D');
```

`convertToRoman(501)` 는 문자열 `DI` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(501), 'DI');
```

`convertToRoman(649)` 는 문자열 `DCXLIX` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(649), 'DCXLIX');
```

`convertToRoman(798)` 는 문자열 `DCCXCVIII` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(798), 'DCCXCVIII');
```

`convertToRoman(891)` 는 문자열 `DCCCXCI` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(891), 'DCCCXCI');
```

`convertToRoman(1000)` 는 문자열 `M` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(1000), 'M');
```

`convertToRoman(1004)` 는 문자열 `MIV` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(1004), 'MIV');
```

`convertToRoman(1006)` 는 문자열 `MVI` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(1006), 'MVI');
```

`convertToRoman(1023)` 는 문자열 `MXXIII` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(1023), 'MXXIII');
```

`convertToRoman(2014)` 는 문자열 `MMXIV` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(2014), 'MMXIV');
```

`convertToRoman(3999)` 는 문자열 `MMMCMXCIX` 을 반환해야 합니다.

```js
assert.deepEqual(convertToRoman(3999), 'MMMCMXCIX');
```

# --seed--

## --seed-contents--

```js
function convertToRoman(num) {
 return num;
}

convertToRoman(36);
```

# --solutions--

```js
function convertToRoman(num) {
  var ref = [['M', 1000], ['CM', 900], ['D', 500], ['CD', 400], ['C', 100], ['XC', 90], ['L', 50], ['XL', 40], ['X', 10], ['IX', 9], ['V', 5], ['IV', 4], ['I', 1]];
  var res = [];
  ref.forEach(function(p) {
    while (num >= p[1]) {
      res.push(p[0]);
      num -= p[1];
    }
  });
  return res.join('');
}
```
