---
id: aff0395860f5d3034dc0bfc9
title: 전화번호 유효성 검사기
challengeType: 5
saveSubmissionToDB: true
forumTopicId: 16090
dashedName: telephone-number-validator
---

# --description--

전달된 문자열이 유효한 미국 전화번호 형식인지 확인하고 `true`를 반환하세요.

사용자는 유효한 미국 번호 형식이라면 어떤 방식으로든 폼 필드를 작성할 수 있습니다. 다음은 미국 번호의 유효한 형식 예시입니다(아래 테스트에서 다른 변형도 참고하세요):

<blockquote>555-555-5555<br>(555)555-5555<br>(555) 555-5555<br>555 555 5555<br>5555555555<br>1 555 555 5555</blockquote>

이 도전 과제에서는 `800-692-7753` 또는 `8oo-six427676;laskdjf`와 같은 문자열이 주어집니다. 위에 제공된 형식 조합에 따라 미국 전화번호를 유효하거나 무효로 판별하세요. 지역 번호는 필수입니다. 국가 코드가 제공되면 국가 코드가 `1`인지 확인해야 합니다. 문자열이 유효한 미국 전화번호이면 `true`를, 그렇지 않으면 `false`를 반환하세요.

# --hints--

`telephoneCheck("555-555-5555")`는 불리언을 반환해야 합니다.

```js
assert(typeof telephoneCheck('555-555-5555') === 'boolean');
```

`telephoneCheck("1 555-555-5555")`는 `true`를 반환해야 합니다.

```js
assert(telephoneCheck('1 555-555-5555') === true);
```

`telephoneCheck("1 (555) 555-5555")`는 `true`를 반환해야 합니다.

```js
assert(telephoneCheck('1 (555) 555-5555') === true);
```

`telephoneCheck("5555555555")`는 `true`를 반환해야 합니다.

```js
assert(telephoneCheck('5555555555') === true);
```

`telephoneCheck("555-555-5555")`는 `true`를 반환해야 합니다.

```js
assert(telephoneCheck('555-555-5555') === true);
```

`telephoneCheck("(555)555-5555")`는 `true`를 반환해야 합니다.

```js
assert(telephoneCheck('(555)555-5555') === true);
```

`telephoneCheck("1(555)555-5555")`는 `true`를 반환해야 합니다.

```js
assert(telephoneCheck('1(555)555-5555') === true);
```

`telephoneCheck("555-5555")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('555-5555') === false);
```

`telephoneCheck("5555555")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('5555555') === false);
```

`telephoneCheck("1 555)555-5555")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('1 555)555-5555') === false);
```

`telephoneCheck("1 555 555 5555")`는 `true`를 반환해야 합니다.

```js
assert(telephoneCheck('1 555 555 5555') === true);
```

`telephoneCheck("1 456 789 4444")`는 `true`를 반환해야 합니다.

```js
assert(telephoneCheck('1 456 789 4444') === true);
```

`telephoneCheck("123**&!!asdf#")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('123**&!!asdf#') === false);
```

`telephoneCheck("55555555")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('55555555') === false);
```

`telephoneCheck("(6054756961)")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('(6054756961)') === false);
```

`telephoneCheck("2 (757) 622-7382")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('2 (757) 622-7382') === false);
```

`telephoneCheck("0 (757) 622-7382")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('0 (757) 622-7382') === false);
```

`telephoneCheck("-1 (757) 622-7382")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('-1 (757) 622-7382') === false);
```

`telephoneCheck("2 757 622-7382")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('2 757 622-7382') === false);
```

`telephoneCheck("10 (757) 622-7382")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('10 (757) 622-7382') === false);
```

`telephoneCheck("27576227382")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('27576227382') === false);
```

`telephoneCheck("(275)76227382")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('(275)76227382') === false);
```

`telephoneCheck("2(757)6227382")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('2(757)6227382') === false);
```

`telephoneCheck("2(757)622-7382")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('2(757)622-7382') === false);
```

`telephoneCheck("555)-555-5555")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('555)-555-5555') === false);
```

`telephoneCheck("(555-555-5555")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('(555-555-5555') === false);
```

`telephoneCheck("(555)5(55?)-5555")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('(555)5(55?)-5555') === false);
```

`telephoneCheck("55 55-55-555-5")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('55 55-55-555-5') === false);
```

`telephoneCheck("11 555-555-5555")`는 `false`를 반환해야 합니다.

```js
assert(telephoneCheck('11 555-555-5555') === false);
```

`telephoneCheck()`는 유효한 번호로 호출되면 `true`를 반환해야 합니다.

```js

const validPatterns = [
  '1 XXX-XXX-XXXX',
  '1 (XXX) XXX-XXXX',
  '1(XXX)XXX-XXXX',
  '1 XXX XXX XXXX',
  'XXXXXXXXXX',
  'XXX-XXX-XXXX',
  '(XXX)XXX-XXXX',
];

validPatterns.forEach(pattern => {
  while (pattern.includes('X')) {
    pattern = pattern.replace('X',  Math.floor(Math.random() * 7) + 2); //While this may seem weird at first, it's required for the CI build to pass
    //This is apparently because the solution provided for CI purposes actually checks for valid area and exchange codes.
  }
  assert.isTrue(telephoneCheck(pattern));
});
```

`telephoneCheck()`는 무효한 번호로 호출되면 `false`를 반환해야 합니다.

```js

const invalidPatterns = [
  '10 XXX-XXX-XXXX',
  '1 (XX)XXX-XXXX',
  '1!(XXX)XXX-XXXX',
  '-1 XXX XXX XXXX',
  'XXXXXXXX',
  'XXX#XXX-XXXX',
  '(XXXXXX-XXXX',
];

invalidPatterns.forEach(pattern => {
  while (pattern.includes('X')) {
    pattern = pattern.replace('X',  Math.floor(Math.random() * 10));
  }
  assert.isFalse(telephoneCheck(pattern));
});
```


# --seed--

## --seed-contents--

```js
function telephoneCheck(str) {
  return true;
}

telephoneCheck("555-555-5555");
```

# --solutions--

```js
var re = /^([+]?1[\s]?)?((?:[(](?:[2-9]1[02-9]|[2-9][02-8][0-9])[)][\s]?)|(?:(?:[2-9]1[02-9]|[2-9][02-8][0-9])[\s.-]?)){1}([2-9]1[02-9]|[2-9][02-9]1|[2-9][02-9]{2}[\s.-]?){1}([0-9]{4}){1}$/;

function telephoneCheck(str) {
  return re.test(str);
}

telephoneCheck("555-555-5555");
```
