---
id: 641da895fb7ec648a5bdf19c
title: 第 15 步
challengeType: 0
dashedName: step-15
---

# --description--

現在你已經擁有想要包含在 UI 中的所有內容。下一步是讓 `Load More Authors` 按鈕在每次被點擊時取得更多作者。你可以透過為按鈕添加 `click` 事件，並小心地遞增 `startingIndex` 和 `endingIndex` 變數來達成這個目的。

使用 `const` 和箭頭函式語法來創建一個 `fetchMoreAuthors` 函式。暫時不要在裡面放任何東西。確保使用大括號，因為函式內會有多個運算式。

# --before-all--

```js
window.fetch = () => Promise.resolve({json: () => Promise.resolve([{ author: 'Whoever', image: 'data:image/gif;base64,R0lGODlhAQABAIAAAAAAAP///ywAAAAAAQABAAACAUwAOw==', url: "http://not-a-real-url.nowhere/", bio: 'words go here' }])});
```    

# --hints--

你應該使用 `const` 來創建 `fetchMoreAuthors` 函式（程式）。

```js
const explorer = await __helpers.Explorer(code);
const { fetchMoreAuthors } = explorer.allFunctions;
assert.isTrue(fetchMoreAuthors?.toString().startsWith('const'));
```

你的 `fetchMoreAuthors` 變數應該是一個函式。

```js
assert.isFunction(fetchMoreAuthors)
```

你的 `fetchMoreAuthors` 函式不應該帶有任何參數。

```js
const explorer = await __helpers.Explorer(code);
const { fetchMoreAuthors } = explorer.allFunctions;
assert.deepEqual(fetchMoreAuthors?.parameters.map((p) => p.toString()), []);
```

你的 `fetchMoreAuthors` 函式應該使用箭頭函式語法並且函式體為空。

```js
const explorer = await __helpers.Explorer(code);
const { fetchMoreAuthors } = explorer.allFunctions;
const validExpressions = [
  'const fetchMoreAuthors = () => {}',
  'let fetchMoreAuthors = () => {}',
  'var fetchMoreAuthors = () => {}'
];
assert.isTrue(
  validExpressions.some((expression) => fetchMoreAuthors?.matches(expression))
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

let startingIndex = 0;
let endingIndex = 8;
let authorDataArr = [];

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    authorDataArr = data;
    displayAuthors(authorDataArr.slice(startingIndex, endingIndex));
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });

--fcc-editable-region--

--fcc-editable-region--

const displayAuthors = (authors) => {
  authors.forEach(({ author, image, url, bio }, index) => {
    authorContainer.innerHTML += `
      <div id="${index}" class="user-card">
        <h2 class="author-name">${author}</h2>
        <img class="user-img" src="${image}" alt="${author} avatar" />
        <p class="bio">${bio}</p>
        <a class="author-link" href="${url}" target="_blank">${author}'s author page</a>
      </div>
    `;
  });
};
```
