---
id: 641da5462576784453146ec2
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

既然你已經擁有所需的資料，就可以用它來填充 UI。 

但是你正在擷取的資料總共有 26 位作者。你應該一次顯示 8 位作者，而不是一次顯示全部 26 位作者。

首先使用 `let` 關鍵字創建兩個名為 `startingIndex` 和 `endingIndex` 的變數，並分別指定它們的數字為 `0` 和 `8`。你需要在這裡使用 `let`，因為你稍後會重新指定這些值。

然後，創建一個 `authorDataArr` 變數並指定一個空陣列。 

在不同的行上宣告並初始化每個變數。

# --before-all--

```js
window.fetch = () => Promise.resolve({json: () => Promise.resolve([{ author: 'Whoever', image: 'data:image/gif;base64,R0lGODlhAQABAIAAAAAAAP///ywAAAAAAQABAAACAUwAOw==', url: "http://not-a-real-url.nowhere/", bio: 'words go here' }])});
```

# --hints--

你應該使用 `let` 來宣告名為 `startingIndex` 的變數。

```js
assert.match(code, /let\s+startingIndex/);
```

你的 `startingIndex` 變數應該是一個數字。

```js
assert.isNumber(startingIndex);
```

你的 `startingIndex` 變數應該設定為 `0`。

```js
assert.strictEqual(startingIndex, 0);
```

你應該使用 `let` 來宣告名為 `endingIndex` 的變數。

```js
assert.match(code, /let\s+endingIndex/);
```

你的 `endingIndex` 變數應該是一個數字。
  
  ```js
assert.isNumber(endingIndex);
```

你應該將你的 `endingIndex` 變數設為 `8`。

```js
assert.strictEqual(endingIndex, 8)
```

你應該有一個名為 `authorDataArr` 的變數。

```js
assert.isDefined(authorDataArr);
```

你的 `authorDataArr` 變數應該是一個空的陣列。

```js
assert.isArray(authorDataArr);
assert.isEmpty(authorDataArr);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

--fcc-editable-region--

--fcc-editable-region--

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    console.log(data);   
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });
```
