---
id: 6606933d6813a8308c962dd1
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

為了成為一個實體方法，方法需要一個特殊的參數，慣例上命名為 `self`。這個參數是對類別實體的引用，且必須永遠是第一個參數。

為你的 `spam` 方法添加一個 `self` 參數。

# --hints--

你的 `spam` 方法應該有一個 `self` 參數。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam():
        pass
    
gameboard = Board()
--fcc-editable-region--
```
