---
id: 6606d8c323d6205890fbbd54
title: 步驟 46
challengeType: 20
dashedName: step-46
---

# --description--

在 `Board` 類別中，創建另一個方法 `is_valid` 並給它三個參數：`self`、`empty` 和 `num`。其中 `empty` 是一個表述空白格列和行索引的元組，`num` 是要檢查的數字。

此方法將透過驗證該數字與指定空白格所在的列、行及 3x3 方格的相容性，來檢查該數字是否為數獨棋盤中空白格的有效選擇。

# --hints--

你應該在 `Board` 類別中創建一個名為 `is_valid` 的新方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("is_valid")`)) })
```

你的 `is_valid` 方法應該有三個參數：`self`、`empty` 和 `num`，且順序如上。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("is_valid").has_args("self, empty, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))

    def valid_in_square(self, row, col, num):
        row_start = (row // 3) * 3
        col_start = (col // 3) * 3
        for row_no in range(row_start, row_start + 3):
            for col_no in range(col_start, col_start + 3):
                if self.board[row_no][col_no] == num:
                    return False
        return True
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
