---
id: 6606b224a69a293f98f8db8f
title: Step 17
challengeType: 20
dashedName: step-25
---

# --description--

當值未被找到時，`.index()` 方法會引起 `ValueError` 異常。 爲了防止程序停止執行，你將要在 `try` 塊中嵌套這行代碼。 `try` 語句用於封裝可能引起異常的代碼。 另一方面，`except` 子句提供了在發生異常時執行的替代代碼：

```py
try:
    <code>
except:
    <code>
```

將 `col` 賦值代碼放到 `try` 代碼塊中。 然後，創建一條 `except` 子句，並在裏面填充 `pass` 。

# --hints--

你應該把 `col` 賦值語句放到 `try` 代碼塊中，並創建 `except` 子句。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_for_loops()[0].find_bodies()[0].is_equivalent("try:\\n  col = contents.index(0)\\nexcept:\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            col = contents.index(0)
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
