---
id: 6606c3fd5634684d48a7887b
title: 步驟 32
challengeType: 20
dashedName: step-32
---

# --description--

你剛剛在前一步撰寫的產生器表達式會產生一個布林值列表，表述指定行中所有列的每個元素是否符合條件 `self.board[row][col] != num` 為 `True` 或 `False`。

將該產生器表達式傳給 `all()` 函式（程式），以檢查該行中的所有元素是否都不同於 `num` 並傳回結果。

# --hints--

你應該將產生器表達式作為引數傳遞給 `all()` 呼叫，並傳回結果。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_col").find_body().is_equivalent("return all(self.board[row][col] != num for row in range(9))")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]
--fcc-editable-region--
    def valid_in_col(self, col, num):
        (self.board[row][col] != num for row in range(9))
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
