---
id: 6606cc754a8834509cd0afb6
title: 第 36 步
challengeType: 20
dashedName: step-36
---

# --description--

另一件要檢查的是數字是否可以插入到 3x3 方格中。

在 `Board` 類別內，創建一個名為 `valid_in_square` 的方法，並帶有四個參數：`self`、`row`、`col` 和 `num`。其中 `row`、`col` 和 `num` 分別表述要檢查的列索引、行索引和數字。

# --hints--

你應該在 `Board` 類別中創建一個名為 `valid_in_square` 的新方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_square")`)) })
```

你的 `valid_in_square` 方法應該有四個參數：`self`、`row`、`col` 和 `num`，且順序如上。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_square").has_args("self, row, col, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
