---
id: 66069b0b36053733a2f012fe
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

現在，刪除你的 `spam` 呼叫。

# --hints--

你的程式碼中不應該有 `gameboard.spam()`。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
gameboard.spam()
--fcc-editable-region--
```
