---
id: 6606927d010be4300a4e5330
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

實例化創建了一個空對象。 但是類可以有方法，它們就像每個實例的本地函數一樣。 在類中，方法聲明如下：

```py
class ClassName:
    def method_name():
        pass
```

在 `Board` 類中，用空方法 `spam` 替換 `pass`。

# --hints--

你應該在 `Board` 類中定義一個名爲 `spam` 的方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass

gameboard = Board()
--fcc-editable-region--
```
