---
id: 6606a219f9efbf38ad496f67
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

你也可以使用點記法來存取實體屬性。

在 `Board` 類別外，初始化 `gameboard` 物件後，使用 `gameboard.board` 存取你的 `gameboard` 物件的 `board` 屬性，並將結果列印到螢幕上。

# --hints--

你應該列印 `gameboard.board`。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.board)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        self.board = board

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
