---
id: 6606b224a69a293f98f8db8f
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

當找不到值時，`.index()` 方法會引發 `ValueError` 例外。為了防止程式停止執行，你會將這行程式碼巢狀放入 `try` 區塊中。`try` 述語用來包裹可能引發例外的程式碼。另一方面，`except` 語句則提供在發生例外時要執行的替代程式碼：

```py
try:
    <code>
except:
    <code>
```

將 `col` 的指派放入 `try` 區段內。然後，創建一個 `except` 語句並在其主體中填入 `pass`。

# --hints--

你應該將 `col` 的指派放在 `try` 區段內，並創建一個 `except` 語句。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_for_loops()[0].find_bodies()[0].is_equivalent("try:\\n  col = contents.index(0)\\nexcept:\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            col = contents.index(0)
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
