---
id: 66068fb0bfddba2b7977eb60
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個專案中，你將透過建置一個數獨解題器來學習類別和物件。

在 Python 中，類別是用來創建物件的藍圖。由類別創建的物件是該類別的實體。你可以使用此語法來創建類別：

```py
class ClassName:
    pass
```

其中 `class` 是定義類別所需的關鍵字，而 `ClassName` 是類別的名稱，慣例上以 *PascalCase* 撰寫。

首先建立一個 `Board` 類別。

# --hints--

你應該創建一個名為 `Board` 的類別。

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
