---
id: 66069167b3307b2f4067b22b
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

透過函式（程式）表示法創建類別的新實體，這涉及在類別名稱後附加一對小括弧。

在類別定義式外，創建 `Board` 類別的實體並將其指定給名為 `gameboard` 的變數。

# --hints--

你應該在 `Board` 類別外宣告變數 `gameboard`。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("gameboard")`)) })
```

你的 `gameboard` 變數應該具有 `Board()` 的值。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
--fcc-editable-region--
```
