---
id: 6606b961ebcf04460f8af76e
title: 第 21 步
challengeType: 20
dashedName: step-21
---

# --description--

透過在 `gameboard` 上呼叫 `find_empty_cell` 方法並列印結果，測試該方法是否正常運作。

請注意，雖然 `find_empty_cell` 定義時有一個參數，但你不可在函式呼叫時傳入引數來給它賦值，因為 `self` 會自動作為你呼叫該方法的物件傳入。

# --hints--

你應該列印 `gameboard.find_empty_cell()`。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.find_empty_cell())")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--

--fcc-editable-region--
```
