---
id: 660699119472f332798860ad
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

現在，將 `pass` 取代為 `print` 呼叫並傳入字串 `'Spam!'`。

# --hints--

你應該刪除 `pass` 並在 `spam` 方法中列印字串 `'Spam!'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").find_body().is_equivalent("print('Spam!')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        pass
    
gameboard = Board()
--fcc-editable-region--
```
