---
id: 6606bc4e5535c0484990ccd5
title: 第 24 步
challengeType: 20
dashedName: step-24
---

# --description--

接下來，你將要處理一個方法，用來檢查給定的數字是否可以插入到指定的數獨棋盤的某一列。

在 `Board` 類別中，創建一個名為 `valid_in_row` 的方法，並給它三個參數：`self`、`row` 和 `num`。其中 `self` 表現類別的實體，`row` 和 `num` 分別是要檢查的列索引和數字。


# --hints--

你應該在 `Board` 類別中創建一個名為 `valid_in_row` 的新方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_row")`)) })
```

你的 `valid_in_row` 方法應該有三個參數：`self`、`row` 和 `num`，且順序為此。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_row").has_args("self, row, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
