---
id: 660699aabc59c532f2d556e5
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

要呼叫實體方法，你需要使用點記法：

```py
instance_name.method_name()
```

其中 `instance_name` 是實體或物件，`method_name` 是你想要呼叫的方法。

呼叫 `gameboard` 物件的 `spam` 方法。

# --hints--

你應該使用 `gameboard.spam()` 來呼叫 `gameboard` 物件的 `spam` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
