---
id: 6606a4641ec48b3a9fe8c2fc
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

內建的 `enumerate` 函式以可迭代物件作為引數，並傳回一個你可以迭代的 enumerate 物件。它會提供計數（預設從零開始）以及來自可迭代物件的值。

```py
iterable = ['a', 'b', 'c']
for i, j in enumerate(iterable):
    print(i, j)
```

上面範例中的迴圈會輸出元組 `0, a`、`1, b` 和 `2, c`。

在 `find_empty_cell` 方法中，將 `pass` 替換為使用 `enumerate()` 函式來迭代數獨棋盤中每一列的 `for` 迴圈。使用 `row` 作為目前列的索引，並使用 `contents` 作為目前列的元素。


# --hints--

你應該將 `pass` 取代為一個 `for` 迴圈，該迴圈迭代 `enumerate(self.board)`。使用 `row` 和 `contents` 作為迴圈變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_body().is_equivalent("for row, contents in enumerate(self.board):\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
