---
id: 6606927d010be4300a4e5330
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

具現化實體會創建一個空的物件。但類別可以有方法，這些方法就像是每個實體的區域函式。在類別中，方法的宣告方式如下：

```py
class ClassName:
    def method_name():
        pass
```

在 `Board` 類別中，將 `pass` 取代為空的 `spam` 方法。

# --hints--

你應該在 `Board` 類別中定義一個名為 `spam` 的方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
    
gameboard = Board()
--fcc-editable-region--
```
