---
id: 66069e5759b800364707988e
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

回到 `__init__` 方法，它需要一個額外的參數來表述要解決的拼圖。

為你的 `__init__` 方法添加第二個名為 `board` 的參數，並透過將 `puzzle` 列表作為引數傳遞給 `gameboard` 的具現體來修正它，就像你在函式呼叫中傳遞引數一樣。

# --hints--

你的 `__init__` 方法應該有兩個參數，順序為：`self` 和 `board`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self, board")`)) })
```

你的 `gameboard` 變數應該具有 `Board(puzzle)` 的值。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board(puzzle)")`)) })
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board()
--fcc-editable-region--
```
