---
id: 66069f86f58f85371d47123e
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

屬性是與物件相關的變數，用來像一般變數一樣儲存資料。

在 `__init__` 方法中，使用 `self.board` 將 `board` 參數（在建立 `Board` 類別的實體時傳入）指定給實體屬性 `board`。

`self.board` 取用該類別實體的 `board` 屬性。它是屬於從 `Board` 類別生成的物件的變數。

# --hints--

你應該刪除 `pass` 並在 `__init__` 方法中將 `board` 參數指定給 `self.board`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").find_body().is_equivalent("self.board = board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
