---
id: 66069b992c1c5e3451f3deb0
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

具現體會創建一個空的物件。`__init__` 方法是一個特殊的方法，允許你將物件具現化為訂製的狀態。當一個類別實作 `__init__` 方法時，`__init__` 會在具現體時自動被呼叫。

在你的 `Board` 類別中，刪除 `spam` 方法，並以包含 `self` 參數的 `__init__` 方法取代它。

# --hints--

你的 `Board` 類別中不應該有 `spam` 方法。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

你應該在你的 `Board` 類別中定義一個 `__init__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("__init__")`)) })
```

你的 `__init__` 方法應該有一個 `self` 參數。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
