---
id: 6606a3ccb1eea93a23c066bf
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

現在你將處理一個尋找數獨棋盤中空白格子的 `method`。

在 `Board` 類別中，創建一個名為 `find_empty_cell` 的空方法，並給它一個 `self` 參數。

# --hints--

你應該在你的 `Board` 類別中定義一個名為 `find_empty_cell` 的方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("find_empty_cell")`)) })
```

你的 `find_empty_cell` 方法應該有一個參數 `self`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        self.board = board
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
