---
id: 6606c14182435d4bab0de2ee
title: 第 30 步
challengeType: 20
dashedName: step-30
---

# --description--

接下來，你將創建一個方法，透過檢查該數字是否尚未出現在指定的數獨棋盤行中，來判斷該數字是否可以插入該行。

在 `Board` 類別中，創建一個名為 `valid_in_col` 的方法，並給它三個參數：`self`、`col` 和 `num`。其中 `col` 和 `num` 分別是要檢查的行索引和數字。

# --hints--

你應該在 `Board` 類別中創建一個名為 `valid_in_col` 的新方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_col")`)) })
```

你的 `valid_in_col` 方法應該有三個參數：`self`、`col` 和 `num`，且順序為此。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_col").has_args("self, col, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
