---
id: 66069b992c1c5e3451f3deb0
title: 步驟 8
challengeType: 20
dashedName: step-8
---

# --description--

實例化創建了一個空對象。 `__init__` 方法允許你實例化一個自定義的對象。 當一個類實現 `__init__` 方法時，實例化時會自動調用 `__init__`。

在你的 `Board` 類中，刪除 `spam` 方法並將其替換爲包含 `self` 參數的 `__init__` 方法。

# --hints--

在 `Board` 類中不應該存在 `spam` 方法。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

你應該在 `Board` 類中定義一個 `__init__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("__init__")`)) })
```

`__init__` 方法應該接收一個 `self` 參數。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')

gameboard = Board()
--fcc-editable-region--
```
