---
id: 0095d0567b6817114989a1dd
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, implementarás el algoritmo de camino más corto en JavaScript. Escribirás una función que calcule el camino más corto entre nodos en un grafo ponderado y devuelva tanto las distancias como las rutas tomadas.

Por ejemplo, dado un grafo donde las ciudades están conectadas por carreteras con diferentes distancias, el algoritmo encontrará la ruta más corta de una ciudad a otra. Si quieres viajar de la ciudad A a la ciudad D, el algoritmo podría encontrar que ir A ⇨ B ⇨ C ⇨ D (total: 15 km) es más corto que ir directamente A ⇨ D (20 km).

Para comenzar, define una variable llamada `INF` y asígnale el valor `Infinity`. Más adelante, la usarás para indicar que no hay conexión directa entre dos nodos.

# --hints--

Debes definir una variable llamada `INF`.

```js
assert.isDefined(INF);
```

Tu variable `INF` debe tener el valor `Infinity`.

```js
assert.strictEqual(INF, Infinity);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
