---
id: 69c918841987bf16c1d1bac8
title: Paso 17
challengeType: 1
dashedName: step-17
---

# --description--

Antes de actualizar las distancias, necesitas verificar que el vecino valga la pena considerarlo. Debe existir una arista real hacia él (`distance !== INF`) y no debe haber sido visitado aún.

Dentro del ciclo `for` de vecinos, añade una sentencia `if` que verifique si `distance` no es estrictamente igual a `INF` **y** si `nodeNo` no ha sido visitado.

Dentro del bloque `if`, declara una variable `newDistance` asignada a `distances[current] + distance`. Este es el costo total de llegar al vecino a través del nodo actual.

# --hints--

Debes añadir una sentencia `if` con la condición `distance !== INF && !visited[nodeNo]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*distance\s*!==\s*INF\s*&&\s*!\s*visited\s*\[\s*\w+\s*\]\s*\)/
);
```

Debes declarar una `newDistance` con el valor `distances[current] + distance` dentro del bloque `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+newDistance\s*=\s*distances\s*\[\s*current\s*\]\s*\+\s*distance/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
