---
id: 69c918841987bf16c1d1babd
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

Necesitas llevar un registro de la distancia más corta conocida desde el nodo inicial hasta cada otro nodo. Para empezar, asumirás que cada nodo está a una distancia infinita.

En JavaScript, puedes crear un arreglo prellenado con un valor usando `new Array(n).fill(value)`. Por ejemplo, `new Array(3).fill(0)` crea `[0, 0, 0]`.

Dentro de la función `shortestPath`, crea una variable llamada `distances` y asígnale `new Array(n).fill(INF)`.

# --hints--

Debes crear una variable llamada `distances` dentro de `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distances\s*=/
);
```

Tu variable `distances` debe estar asignada a `new Array(n).fill(INF)`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*INF\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;

}
--fcc-editable-region--
```
