---
id: 69c918841987bf16c1d1bac1
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

Ahora agregarás el ciclo principal que impulsa el algoritmo. Se ejecuta una vez por cada nodo en el grafo, seleccionando el nodo no visitado más cercano cada vez.

Dentro de la función `shortestPath`, agrega un ciclo `for` que se ejecute `n` veces usando `i` como variable del ciclo.

Al inicio del cuerpo del ciclo, usando `let`, declara dos variables `minDistance` y `current` asignadas a `INF` y `-1` respectivamente. Estas rastrearán el nodo no visitado más cercano encontrado en cada iteración.

# --hints--

Deberías tener un ciclo `for` dentro de `shortestPath` que se ejecute `n` veces. No olvides asignar `i` con `let` dentro de él.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+i\s*=\s*0\s*;\s*i\s*<\s*n\s*;\s*i\+\+\s*\)/
);
```

Deberías tener una variable `minDistance` asignada a `INF` dentro del ciclo.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+minDistance\s*=\s*INF/
);
```

Deberías tener una variable `current` asignada a `-1` dentro del ciclo.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+current\s*=\s*-1/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
--fcc-editable-region--

--fcc-editable-region--
}
```
