---
id: 69c918841987bf16c1d1bac3
title: Paso 12
challengeType: 1
dashedName: step-12
---

# --description--

Dentro del bucle interno `for`, necesitas verificar si el nodo actual no ha sido visitado y está más cerca que el mejor que has encontrado hasta ahora.

Agrega una sentencia `if` que compruebe si `nodeNo` no ha sido visitado y `distances[nodeNo]` es menor que `minDistance`.

Deja el cuerpo vacío por ahora.

# --hints--

Deberías tener una sentencia `if` con la condición `!visited[nodeNo] && distances[nodeNo] < minDistance` dentro del bucle interno `for`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*!visited\s*\[\s*nodeNo\s*\]\s*&&\s*distances\s*\[\s*nodeNo\s*\]\s*<\s*minDistance\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
