---
id: 69c918841987bf16c1d1babf
title: Paso 8
challengeType: 1
dashedName: step-8
---

# --description--

También necesitas rastrear la ruta real tomada para llegar a cada nodo. Guardarás una lista de rutas donde cada entrada es un arreglo de índices de nodos que representan la ruta seguida.

Inicialmente, la ruta de cada nodo solo contendrá a sí mismo. Puedes crear esta estructura usando `Array.from()`:

```js
Array.from({ length: n }, (_, i) => [i]);
```

Esto crea un arreglo de `n` elementos, donde el elemento en el índice `i` es `[i]`.

Aún dentro de la función `shortestPath`, crea una variable llamada `paths` y asígnale `Array.from({ length: n }, (_, i) => [i])`.

# --hints--

Debes crear una variable `paths` dentro de la función `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+paths\s*=/
);
```

Tu variable `paths` debe inicializarse usando `Array.from`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*=\s*Array\.from/
);
```

Tu variable `paths` debe usar `{ length: n }` como primer argumento de `Array.from`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\(\s*\{\s*length\s*:\s*n\s*\}/
);
```

El segundo argumento de `Array.from` debe ser un callback `(_, i) => [i]`, donde `_` es un marcador para el elemento no usado y `i` es el índice. Debe devolver `[i]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\([\s\S]*?,\s*function\s*\(\s*\w+\s*,\s*(\w+)\s*\)\s*\{[\s\S]*?return\s*\[\s*\1\s*\]|Array\.from\s*\([\s\S]*?,\s*\(\s*\w+\s*,\s*\w+\s*\)\s*=>\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
