---
id: 69c918841987bf16c1d1bac9
title: Paso 18
challengeType: 1
dashedName: step-18
---

# --description--

Ahora deberías verificar si pasar por el nodo actual ofrece una ruta más corta hacia el vecino. Si `newDistance` es mejor que lo que ya está almacenado, entonces debes actualizarlo.

Dentro del bloque `if` existente, añade una sentencia `if` anidada que compruebe si `newDistance` es menor que `distances[nodeNo]`. Dentro de ese bloque `if` anidado, actualiza `distances[nodeNo]` a `newDistance`.

# --hints--

Debes añadir una sentencia `if` anidada con la condición `newDistance < distances[nodeNo]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*newDistance\s*<\s*distances\s*\[\s*\w+\s*\]\s*\)/
);
```

Debes actualizar `distances[nodeNo]` a `newDistance` dentro del bloque `if` anidado.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /distances\s*\[\s*\w+\s*\]\s*=\s*newDistance/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        --fcc-editable-region--

        --fcc-editable-region--
      }
    }
  }
}
```
