---
id: 69c918841987bf16c1d1baba
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Necesitarás un arreglo 2D para representar la matriz de adyacencia del grafo. Cada valor representa el peso (distancia) de la arista entre dos nodos. Un valor de `INF` significa que no hay una arista directa.

Crea otra variable llamada `adjMatrix` y asígnale un arreglo 2D con los siguientes valores:

```js
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

Debes crear una variable llamada `adjMatrix`.

```js
assert.isDefined(adjMatrix);
```

`adjMatrix` debe ser un arreglo 2D con 6 filas y 6 columnas.

```js
assert.isArray(adjMatrix);
assert.lengthOf(adjMatrix, 6);
adjMatrix.forEach(row => assert.lengthOf(row, 6));
```

`adjMatrix` debe contener los valores correctos.

```js
assert.deepEqual(adjMatrix[0], [0, 5, 3, INF, 11, INF]);
assert.deepEqual(adjMatrix[1], [5, 0, 1, INF, INF, 2]);
assert.deepEqual(adjMatrix[2], [3, 1, 0, 1, 5, INF]);
assert.deepEqual(adjMatrix[3], [INF, INF, 1, 0, 9, 3]);
assert.deepEqual(adjMatrix[4], [11, INF, 5, 9, 0, INF]);
assert.deepEqual(adjMatrix[5], [INF, 2, INF, 3, INF, 0]);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
--fcc-editable-region--

--fcc-editable-region--
```
