---
id: 69c918841987bf16c1d1bac7
title: Paso 16
challengeType: 1
dashedName: step-16
---

# --description--

Ahora que el nodo actual está marcado como visitado, necesitas revisar todos sus vecinos para ver si puedes alcanzarlos de manera más eficiente a través del nodo actual.

Después de `visited[current] = true`, agrega un bucle vecino `for` que itere por cada `nodeNo` desde `0` hasta `n`. Ten en cuenta que este es otro bucle interno al mismo nivel que el primero.

Dentro del bucle, declara una variable llamada `distance` y asígnale `matrix[current][nodeNo]`. Esto te da el peso de la arista entre el nodo actual y su vecino.

# --hints--

Debes agregar un bucle vecino `for` después de marcar el nodo actual como visitado.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(/
);
```

El bucle vecino `for` debe iterar desde `0` hasta `n`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(\s*(var|let)\s+\w+\s*=\s*0\s*;\s*\w+\s*<\s*n\s*;\s*\w+\+\+\s*\)/
);
```

Debes tener una variable `distance` asignada a `matrix[current][nodeNo]` dentro del bucle vecino.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distance\s*=\s*matrix\s*\[\s*current\s*\]\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
