---
id: 69c918841987bf16c1d1babc
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Necesitas almacenar el número de nodos en el grafo.

Dentro de la función `shortestPath`, crea una variable llamada `n` y asígnale `matrix.length`.

# --hints--

Debes crear una variable `n` dentro de la función `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+n\s*=/
);
```

Tu variable `n` debe tener el valor `matrix.length`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+n\s*=\s*matrix\.length/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {

}
--fcc-editable-region--
```
