---
id: 69c918841987bf16c1d1bac4
title: Paso 13
challengeType: 1
dashedName: step-13
---

# --description--

Si la condición es verdadera, el nodo actual es el mejor candidato no visitado encontrado hasta ahora. Necesitas actualizar tus variables de seguimiento para reflejar eso.

Dentro de la sentencia `if`, actualiza `minDistance` a `distances[nodeNo]` y `current` a `nodeNo`.

# --hints--

Debes asignar `minDistance` a `distances[nodeNo]` dentro del bloque `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /minDistance\s*=\s*distances\s*\[\s*nodeNo\s*\]/
);
```

Debes asignar `current` a `nodeNo` dentro del bloque `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /current\s*=\s*nodeNo/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        --fcc-editable-region--

        --fcc-editable-region--
      }
    }
  }
}
```
