---
id: 6143869bb45bd85e3b1926aa
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Dentro de tu elemento `.heading`, crea un elemento `header` con la `clase` establecida en `hero`.

En ese elemento, crea un elemento `img` con el `src` establecido en `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`, el `alt` establecido en `freecodecamp logo`, y la `clase` establecida en `hero-img`.

El atributo `loading` en un elemento `img` se puede establecer en `lazy` para indicarle al navegador que no busque el recurso de imagen hasta que sea necesario (es decir, cuando el usuario desplace la imagen en la vista). Como beneficio adicional, los elementos cargados de manera diferida no se cargarán hasta que los elementos no diferidos estén cargados; esto significa que los usuarios con conexiones a internet lentas pueden ver el contenido de tu página sin tener que esperar a que las imágenes se carguen.

Da a tu nuevo elemento `img` un atributo `loading` establecido en `lazy`.

# --hints--

Debes crear un elemento `header`.

```js
assert.exists(document.querySelector('header'));
```

Tu elemento `header` debe estar dentro de tu elemento `.heading`.

```js
assert.equal(document.querySelector('header')?.parentElement?.className, 'heading');
```

Tu elemento `header` debe tener la `clase` establecida en `hero`.

```js
assert.equal(document.querySelector('header')?.className, 'hero');
```

Debes crear un elemento `img`.

```js
assert.exists(document.querySelector('img'));
```

Tu elemento `img` debe estar dentro de tu elemento `header`.

```js
assert.equal(document.querySelector('img')?.parentElement?.localName, 'header');
```

Tu elemento `img` debe tener el `src` establecido en `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`.

```js
assert.equal(document.querySelector('img')?.getAttribute('src'), 'https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png');
```

Tu elemento `img` debe tener el `alt` establecido en `freecodecamp logo`.

```js
assert.equal(document.querySelector('img')?.getAttribute('alt'), 'freecodecamp logo');
```

Tu elemento `img` debe tener el atributo `loading` establecido en `lazy`.

```js
assert.equal(document.querySelector('img')?.getAttribute('loading'), 'lazy');
```

Tu elemento `img` debe tener la `clase` establecida en `hero-img`.

```js
assert.equal(document.querySelector('img')?.className, 'hero-img');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
