---
id: 614385513d91ae5c251c2052
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, crearás una página de revista usando un comunicado de prensa de una versión anterior del currículo de freeCodeCamp.

Comienza con añadir un elemento `link` para la hoja de estilos de fuente `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`, un `link` para la hoja de estilos de FontAwesome `https://use.fontawesome.com/releases/v5.8.2/css/all.css`, y un `link` para tu hoja de estilos `./styles.css`.

La hoja de estilo de fuente cargará tres fuentes separadas: `Anton`, `Baskervville` y `Raleway`.

# --hints--

El código debe tener tres elementos `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

Los elementos `link` deben estar dentro del elemento `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

Tus tres elementos `link` deben tener un atributo `rel` con el valor `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Uno de los elementos del enlace debe tener el `href` establecido en`https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

Uno de los elementos del enlace debe tener el `href` establecido en`https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

Uno de los elementos `link` debe tener el atributo `href` con un valor de `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
