---
id: 641da5462576784453146ec2
title: 步驟7
challengeType: 0
dashedName: step-7
---

# --description--

現在你有了所需的數據，可以使用它來填充 UI。 但是獲取的數據包含 26 位作者的數組，如果將他們全部同時添加到頁面，可能會導致性能不佳。

你應該一次添加 8 位作者，然後有一個按鈕來添加另外 8 位作者，直到沒有更多數據可顯示。

使用 `let` 創建 2 個名爲 `startingIndex` 和 `endingIndex` 的變量，並分別向它們分配數字值 `0` 和 `8`。 另外，使用 `let` 創建一個 `authorDataArr` 變量並將其設置爲一個空數組。

# --before-all--

```js
window.fetch = () => Promise.resolve({json: () => Promise.resolve([{ author: 'Whoever', image: 'http://not-a-real-url.nowhere/no-image.jpg', url: "http://not-a-real-url.nowhere/", bio: 'words go here' }])});
```

# --hints--

你應該使用 `let` 來聲明一個名爲 `startingIndex` 的變量。

```js
assert.match(code, /let\s+startingIndex/)
```

你應該將 `startingIndex` 變量設置爲 `0`。

```js
assert.match(code, /let\s+startingIndex\s*=\s*0\s*;?/)
```

你應該使用 `let` 來聲明一個名爲 `endingIndex` 的變量。

```js
assert.match(code, /let\s+endingIndex/)
```

你應該將 `endingIndex` 變量設置爲 `8`。

```js
assert.match(code, /let\s+endingIndex\s*=\s*8\s*;?/)
```

你應該使用 `let` 來聲明一個名爲 `authorDataArr` 的變量。

```js
assert.match(code, /let\s+authorDataArr/)
```

你應該將 `authorDataArr` 變量設置爲空數組 (`[]`)。

```js
assert.match(code, /let\s+authorDataArr\s*=\s*\[\s*\]\s*;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

--fcc-editable-region--

--fcc-editable-region--

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    console.log(data);   
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });
```
