---
id: 68de40f53fcaf106f51cc20f
title: HTMLのオーディオとビデオプレーヤーを作成する
challengeType: 25
dashedName: build-an-html-audio-and-video-player
demoType: onClick
---

# --description--

**目的：** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー：**

1. ページのメインタイトルには`h1`要素が必要です。
2. `section`要素を2つ用意してください。
3. 最初の`section`要素の中に、再生中のビデオのタイトル用の`h2`要素が必要です。
4. `h2`要素の下に、`video`および`controls`属性を持つ`width`要素が必要です。`width`属性は`640`に設定してください。
5. `video`要素の中に、ビデオファイルを指す`source`属性と`src`属性を持つ`type`要素が必要です。
   - `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`を使うことができます。
6. 2番目の`section`要素の中に、再生中の曲のタイトル用の`h2`要素が必要です。
7. `h2`要素の下に、`audio`および`controls`属性を持ち、オーディオファイルを指す`loop`属性を持つ`src`要素が必要です。
   - `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`を使うことができます。
   - または`https://cdn.freecodecamp.org/curriculum/js-music-player/we-are-going-to-make-it.mp3`を使うことができます。

# --hints--

ページのメインタイトルには`h1`要素が必要です。

```js
const h1 = document.querySelector('h1');
assert.exists(h1);
```

`h1`は空にしてはいけません。

```js
const h1 = document.querySelector('h1');
assert.isNotEmpty(h1.textContent.trim());
```

`h1`要素は1つだけにしてください。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

`section`要素は2つ必要です。

```js
assert.lengthOf(document.querySelectorAll('section'), 2);

```

最初の`section`要素の中に、再生中のビデオのタイトル用の`h2`要素が必要です。

```js
const section = document.querySelector('section');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

`h2`要素の下に`video`要素が必要です。

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
assert.exists(video);
```

`video`要素は`controls`属性を持つ必要があります。

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const controlsAttr = video.controls;
assert.isTrue(controlsAttr);
```

`video`要素は`width`属性を`640`に設定してください。

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const widthAttr = video.getAttribute('width');
assert.equal(widthAttr, '640');
```

`video`要素の中に`source`要素が必要です。

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
assert.exists(source);
```

`source`要素はビデオファイルを指す`src`属性を持つ必要があります。

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const srcAttr = source.getAttribute('src');
assert.match(srcAttr, /\.(mp4|webm|ogg|avi|mov)$/i);
```

`source`要素は`type`属性を持つ必要があります。

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
assert.isNotEmpty(typeAttr);
```

`type`属性は`src`属性のファイル拡張子のメディアタイプと一致している必要があります。

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
const srcAttr = source.getAttribute('src');
const srcFileExt = srcAttr.match(/[^\.]+$/g)[0];
const fileExtToMediaTypes = {
  mp4: ['video/mp4'],
  webm: ['video/webm'],
  ogg: ['video/ogg'],
  avi: ['video/avi', 'video/vnd.avi', 'video/msvideo', 'video/x-msvideo'],
  mov: ['video/quicktime']
};
const srcMediaTypes = fileExtToMediaTypes[srcFileExt];
assert.oneOf(typeAttr, srcMediaTypes);
```

2番目の`section`要素の中に、再生中の曲のタイトル用の`h2`要素が必要です。

```js
const section = document.querySelector('section:nth-of-type(2)');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

すべての`h2`要素にはテキストが含まれている必要があります。

```js
const h2s = document.querySelectorAll('h2');
assert.isNotEmpty(h2s);
h2s.forEach((h2) => assert.isNotEmpty(h2.textContent.trim()));
```

`h2`要素の下に`audio`要素が必要です。

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
assert.exists(audio);
```

`audio`要素は`controls`属性を持つ必要があります。

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const controlsAttr = audio.controls;
assert.isTrue(controlsAttr);
```

`audio`要素は`loop`属性を持つ必要があります。

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const loopAttr = audio.loop;
assert.isTrue(loopAttr);
```

`audio`要素はオーディオファイルを指す`src`属性を持つ必要があります。

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const srcAttr = audio.getAttribute('src');
assert.match(srcAttr, /\.(mp3|wav|ogg|aac|flac)$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>
  <main>
    <h1>HTML Audio and Video Lab</h1>

    <section>
      <h2>What is the map method and how does it work?</h2>
      <video controls width="640">
        <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4" />
      </video>
    </section>

    <section>
      <h2>Sailing Away</h2>
      <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" controls loop>
      </audio>
    </section>

  </main>
</body>
</html>
```
