---
id: 6823f9df49cc206af5471a30
title: Camperbots Profilseite debuggen
challengeType: 25
dashedName: lab-debug-camperbots-profile-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

Camperbot versucht, eine Profilseite zu erstellen. Er hat einen Freund gebeten, seinen Code zu überprüfen, und dieser sagte, dass einige Fehler darin sind.

Ihre Aufgabe ist es, alle Fehler von Camperbot zu beheben, damit er mit dem Erstellen seiner Profilseite fortfahren kann. Erfüllen Sie die Punkte in den User Stories unten und klicken Sie auf „Check Your Code“, um zu sehen, ob Sie alle Fehler behoben haben.

**User Stories:**

1. Camperbot versucht, ein `heading2`-Element zu verwenden, aber dieses Element existiert nicht. Korrigieren Sie diese Tags, damit ein korrektes Überschrift-Element der zweiten Ebene verwendet wird.
2. Camperbot versucht, zwei Absätze mit `pp` hinzuzufügen, aber diese existieren ebenfalls nicht. Korrigieren Sie sie, damit korrekte Absatz-Tags verwendet werden.
3. Camperbot verwendet ein `h3`-Element für die `Background and Interests`-Zwischenüberschrift, aber es gibt einen Syntaxfehler. Finden Sie das Problem und beheben Sie es.

# --hints--

Sie sollten genau ein `h1`-Element auf der Seite haben.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

Sie sollten das `<h1>Hello from Camperbot!</h1>`-Element nicht von der Seite entfernen.

```js
assert.match(code, /\<h1\s*\>\s*Hello\s*from\s*Camperbot\s*!\s*\<\/h1\s*\>/i);
```

Sie sollten genau ein `h2`-Element auf der Seite haben.

```js
assert.lengthOf(document.querySelectorAll("h2"), 1);
```

Sie sollten ein `h2`-Element mit dem Text `About` haben. Hier ist ein Beispiel: `<h2>Text here</h2>`

```js
assert.match(code, /\<h2\s*\>\s*About\s*\<\/h2\s*\>/i);
```

Sie sollten ein Absatz-Element mit dem Text `My name is Camperbot and I love learning new things.` haben. Hier ist ein Beispiel: `<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*My\s*name\s*is\s*camperbot\s*and\s*I\s*love\s*learning\s*new\s*things\s*\.\s*\<\/p\s*\>/i);
```

Sie sollten ein Absatz-Element mit dem Text `I enjoy solving puzzles.` haben. Hier ist ein Beispiel: `<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*I\s*enjoy\s*solving\s*puzzles\s*\.\s*\<\/p\s*\>/i);
```

Sie sollten genau zwei Absatz-Elemente auf der Seite haben.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

Sie sollten genau ein `h3`-Element auf der Seite haben.

```js
assert.lengthOf(document.querySelectorAll("h3"), 1);
```

Sie sollten ein `h3`-Element mit dem Text `Background and Interests` haben. Hier ist ein Beispiel: `<h3>Text here</h3>`

```js
assert.match(code, /\<h3\s*\>\s*Background\s*and\s*Interests\s*\<\/h3\s*\>/i);
```

Sie sollten insgesamt genau fünf Elemente auf der Seite haben.

```js
assert.lengthOf(document.querySelectorAll("body *:not(script)"), 5);
```

# --seed--

## --seed-contents--

```html
<h1>Hello from Camperbot!</h1>

<heading2>About</heading2>

<pp>My name is Camperbot and I love learning new things.</pp>

<h3>Background and Interests<h3/>
<pp>I enjoy solving puzzles.</pp>
```

# --solutions--

```html
<h1>Hello from Camperbot!</h1>

<h2>About</h2>

<p>My name is Camperbot and I love learning new things.</p>

<h3>Background and Interests</h3>
<p>I enjoy solving puzzles.</p>
```
