---
id: 587d8249367417b2b2512c41
title: メートル・インペリアル変換器
challengeType: 4
forumTopicId: 301570
dashedName: metric-imperial-converter
---

# --description--

このリンク先の機能と同様のフルスタックJavaScriptアプリを作成してください：<a href="https://metric-imperial-converter.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://metric-imperial-converter.freecodecamp.rocks/</a>。このプロジェクトに取り組むには、以下のいずれかの方法でコードを書いてください。

- <a href="https://github.com/freeCodeCamp/boilerplate-project-metricimpconverter/" target="_blank" rel="noopener noreferrer nofollow">このGitHubリポジトリ</a>をクローンし、ローカルでプロジェクトを完成させる。
- お好きなサイトビルダーを使ってプロジェクトを完成させる。ただし、GitHubリポジトリのすべてのファイルを取り込むこと。

**注意：** このプロジェクトのテストは`glitch.com`を使用すると動作しません。

# --instructions--

- `/controllers/convertHandler.js`に必要な変換ロジックを完成させてください。
- `/routes/api.js`に必要なルートを完成させてください。
- `sample.env`ファイルを`.env`にコピーし、変数を適切に設定してください。
- テストを自動で実行するには、`NODE_ENV=test`ファイルに`.env`を追加してください。
- コンソールでテストを実行するには、コマンド`npm run test`を使用してください。

`tests/1_unit-tests.js`に以下のテストを書いてください：

- `convertHandler`は整数の入力を正しく読み取る必要があります。
- `convertHandler`は小数の入力を正しく読み取る必要があります。
- `convertHandler`は分数の入力を正しく読み取る必要があります。
- `convertHandler`は小数を含む分数の入力を正しく読み取る必要があります。
- `convertHandler`は二重分数（例：`3/2/3`）に対して正しくエラーを返す必要があります。
- `convertHandler`は数値入力がない場合、`1`をデフォルトの数値入力として正しく扱う必要があります。
- `convertHandler`は有効な各入力単位を正しく読み取る必要があります。
- `convertHandler`は無効な入力単位に対して正しくエラーを返す必要があります。
- `convertHandler`は有効な各入力単位に対して正しい返却単位を返す必要があります。
- `convertHandler`は有効な各入力単位に対して単位の綴りを正しく返す必要があります。
- `convertHandler`は`gal`を`L`に正しく変換する必要があります。
- `convertHandler`は`L`を`gal`に正しく変換する必要があります。
- `convertHandler`は`mi`を`km`に正しく変換する必要があります。
- `convertHandler`は`km`を`mi`に正しく変換する必要があります。
- `convertHandler`は`lbs`を`kg`に正しく変換する必要があります。
- `convertHandler`は`kg`を`lbs`に正しく変換する必要があります。

`tests/2_functional-tests.js`に以下のテストを書いてください：

- `10L`のような有効な入力を変換する：`GET`リクエストを`/api/convert`に変換する。
- `32g`のような無効な入力を変換する：`GET`リクエストを`/api/convert`に変換する。
- `3/7.2/4kg`のような無効な数値を変換する：`GET`リクエストを`/api/convert`に変換する。
- `3/7.2/4kilomegagram`のような無効な数値と単位の両方を変換する：`GET`リクエストを`/api/convert`に変換する。
- `kg`のように数値なしで変換する：`GET`リクエストを`/api/convert`に変換する。

# --hints--

独自のプロジェクトを提供しても構いません。例のURLを使う必要はありません。

```js
  assert(
    !/.*\/metric-imperial-converter\.freecodecamp\.rocks/.test(
      code
    )
  );
```

`GET` `/api/convert`を使って、受け入れ可能な数値と単位を含む単一のパラメータを渡し、変換できます。（ヒント：単位の開始を示す最初の文字のインデックスを探して入力を分割してください）

```js

```

`'gal'`を`'L'`に、またその逆も変換できます。（1 gal は 3.78541 L）

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1gal');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.returnNum, 3.78541);
    assert.equal(data1.returnUnit, 'L');
    const response2 = await fetch(code + '/api/convert?input=10gal');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.returnNum, 37.8541);
    assert.equal(data2.returnUnit, 'L');
    const response3 = await fetch(code + '/api/convert?input=1l');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.returnNum, 0.26417);
    assert.equal(data3.returnUnit, 'gal');
    const response4 = await fetch(code + '/api/convert?input=10l');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.returnNum, 2.64172);
    assert.equal(data4.returnUnit, 'gal');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

`'lbs'`を`'kg'`に、またその逆も変換できます。（1 lbs は 0.453592 kg）

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1lbs');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.returnNum, 0.45359);
    assert.equal(data1.returnUnit, 'kg');
    const response2 = await fetch(code + '/api/convert?input=10lbs');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.returnNum, 4.53592);
    assert.equal(data2.returnUnit, 'kg');
    const response3 = await fetch(code + '/api/convert?input=1kg');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.returnNum, 2.20462);
    assert.equal(data3.returnUnit, 'lbs');
    const response4 = await fetch(code + '/api/convert?input=10kg');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.returnNum, 22.04624);
    assert.equal(data4.returnUnit, 'lbs');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

`'mi'`を`'km'`に、またその逆も変換できます。（1 mi は 1.60934 km）

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1mi');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.returnNum, 1.60934);
    assert.equal(data1.returnUnit, 'km');
    const response2 = await fetch(code + '/api/convert?input=10mi');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.returnNum, 16.0934);
    assert.equal(data2.returnUnit, 'km');
    const response3 = await fetch(code + '/api/convert?input=1km');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.returnNum, 0.62137);
    assert.equal(data3.returnUnit, 'mi');
    const response4 = await fetch(code + '/api/convert?input=10km');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.returnNum, 6.21373);
    assert.equal(data4.returnUnit, 'mi');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

すべての入力単位は大文字・小文字の両方で受け入れられますが、返却される単位は`initUnit`および`returnUnit`で小文字に統一されます。ただしリットルは大文字の`'L'`で表されます。

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1gal');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.initUnit, 'gal');
    assert.equal(data1.returnUnit, 'L');
    const response2 = await fetch(code + '/api/convert?input=10L');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.initUnit, 'L');
    assert.equal(data2.returnUnit, 'gal');
    const response3 = await fetch(code + '/api/convert?input=1l');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.initUnit, 'L');
    assert.equal(data3.returnUnit, 'gal');
    const response4 = await fetch(code + '/api/convert?input=10KM');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.initUnit, 'km');
    assert.equal(data4.returnUnit, 'mi');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

単位が無効な場合、`'invalid unit'`が返されます。

```js
  try {
    const response = await fetch(code + '/api/convert?input=1min');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.text();
    assert(data.error === 'invalid unit' || data === 'invalid unit');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

数値が無効な場合、`'invalid number'`が返されます。

```js
  try {
    const response = await fetch(
      code + '/api/convert?input=1//2gal'
    );
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.text();
    assert(data.error === 'invalid number' || data === 'invalid number');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

単位と数値の両方が無効な場合、`'invalid number and unit'`が返されます。

```js
  try {
    const response = await fetch(
      code + '/api/convert?input=1//2min'
    );
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.text();
    assert(
      data.error === 'invalid number and unit' ||
        data === 'invalid number and unit'
    );
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

パラメータには分数、小数、またはその両方を使えます（例：5、1/2、2.5/6）。ただし何も指定されなければデフォルトは1になります。

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=mi');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.approximately(data1.initNum, 1, 0.001);
    assert.approximately(data1.returnNum, 1.60934, 0.001);
    assert.equal(data1.returnUnit, 'km');
    const response2 = await fetch(code + '/api/convert?input=1/5mi');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.approximately(data2.initNum, 1 / 5, 0.1);
    assert.approximately(data2.returnNum, 0.32187, 0.001);
    assert.equal(data2.returnUnit, 'km');
    const response3 = await fetch(
      code + '/api/convert?input=1.5/7km'
    );
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.approximately(data3.initNum, 1.5 / 7, 0.001);
    assert.approximately(data3.returnNum, 0.13315, 0.001);
    assert.equal(data3.returnUnit, 'mi');
    const response4 = await fetch(
      code + '/api/convert?input=3/2.7km'
    );
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.approximately(data4.initNum, 3 / 2.7, 0.001);
    assert.approximately(data4.returnNum, 0.69041, 0.001);
    assert.equal(data4.returnUnit, 'mi');
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

返却値は`initNum`、`initUnit`、`returnNum`、`returnUnit`、`string`で構成され、`'{initNum} {initUnitString} converts to {returnNum} {returnUnitString}'`の形式で単位を綴り、結果は小数点以下5桁に丸められます。

```js
  try {
    const response = await fetch(code + '/api/convert?input=2mi');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.json();
    assert.equal(data.initNum, 2);
    assert.equal(data.initUnit, 'mi');
    assert.approximately(data.returnNum, 3.21868, 0.001);
    assert.equal(data.returnUnit, 'km', 'returnUnit did not match');
    assert.equal(data.string, '2 miles converts to 3.21868 kilometers');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

16個の単位テストはすべて完了し、合格しています。

```js
  try {
    const response = await fetch(code + '/_api/get-tests');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const getTests = await response.json();
    assert.isArray(getTests);
    const unitTests = getTests.filter(test => {
      return !!test.context.match(/Unit Tests/gi);
    });
    assert.isAtLeast(unitTests.length, 16, 'At least 16 tests passed');
    unitTests.forEach(test => {
      assert.equal(test.state, 'passed', 'Tests in Passed State');
      assert.isAtLeast(
        test.assertions.length,
        1,
        'At least one assertion per test'
      );
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

5個の機能テストはすべて完了し、合格しています。

```js
  try {
    const response = await fetch(code + '/_api/get-tests');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const getTests = await response.json();
    assert.isArray(getTests);
    const functTests = getTests.filter(test => {
      return !!test.context.match(/Functional Tests/gi);
    });
    assert.isAtLeast(functTests.length, 5, 'At least 5 tests passed');
    functTests.forEach(test => {
      assert.equal(test.state, 'passed', 'Tests in Passed State');
      assert.isAtLeast(
        test.assertions.length,
        1,
        'At least one assertion per test'
      );
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```
