---
id: 587d824a367417b2b2512c43
title: パーソナルライブラリ
challengeType: 4
forumTopicId: 301571
dashedName: personal-library
---

# --description--

このリンク先のアプリと機能的に似たフルスタックJavaScriptアプリを作成してください：<a href="https://personal-library.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://personal-library.freecodecamp.rocks/</a>。このプロジェクトに取り組むには、以下のいずれかの方法でコードを書いてください。

- <a href="https://github.com/freeCodeCamp/boilerplate-project-library" target="_blank" rel="noopener noreferrer nofollow">このGitHubリポジトリ</a>をクローンして、ローカルでプロジェクトを完成させる。
- お好みのサイトビルダーを使ってプロジェクトを完成させる。必ずGitHubリポジトリのすべてのファイルを組み込んでください。

# --instructions--

1.  MongoDBの接続文字列を`.env`に引用符なしで`DB`として追加してください。
    例：`DB=mongodb://admin:pass@1234.mlab.com:1234/fccpersonallib`
2.  `.env`ファイル内で`NODE_ENV`を引用符なしで`test`に設定してください。
3.  `routes/api.js`内にすべてのルートを作成する必要があります。
4.  `tests/2_functional-tests.js`で全ての機能テストを作成してください。

# --hints--

例のURLではなく、ご自身のプロジェクトを提供してもかまいません。

```js
  assert(
    !/.*\/personal-library\.freecodecamp\.rocks/.test(code)
  );
```

`/api/books`に対してフォームデータの一部として`title`を含む<b>POST</b>リクエストを送信して、本を追加できます。返されるレスポンスは`title`と一意の`_id`をキーとするオブジェクトです。リクエストに`title`が含まれていない場合、返されるレスポンスは文字列`missing required field title`になります。

```js
  try {
    const response1 = await fetch(code + '/api/books', {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ title: 'Faux Book 1' })
    });
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    let data1 = await response1.json();
    assert.isObject(data1);
    assert.property(data1, 'title');
    assert.equal(data1.title, 'Faux Book 1');
    assert.property(data1, '_id');
    const response2 = await fetch(code + '/api/books', {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({})
    });
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    let data2 = await response2.text();
    assert.isString(data2);
    assert.equal(data2, 'missing required field title');
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

`/api/books`に<b>GET</b>リクエストを送信すると、すべての本を表すJSONレスポンスが返されます。JSONレスポンスはオブジェクトの配列で、各オブジェクト（本）は`title`、`_id`、`commentcount`のプロパティを含みます。

```js
  try {
    let url = code + '/api/books';
    let a = fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ title: 'Faux Book A' })
    }).then(response => {
      if (!response.ok) throw Error(response.text());
      return response.json();
    });
    let b = fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ title: 'Faux Book B' })
    }).then(response => {
      if (!response.ok) throw Error(response.text());
      return response.json();
    });
    let c = fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ title: 'Faux Book C' })
    }).then(response => {
      if (!response.ok) throw Error(response.text());
      return response.json();
    });
    await Promise.all([a, b, c]).then(async () => {
      const response = await fetch(url);
      if (!response.ok) {
        throw Error(await response.text());
      }
      let data = await response.json();
      assert.isArray(data);
      assert.isAtLeast(data.length, 3);
      data.forEach((book) => {
        assert.isObject(book);
        assert.property(book, 'title');
        assert.isString(book.title);
        assert.property(book, '_id');
        assert.property(book, 'commentcount');
        assert.isNumber(book.commentcount);
      });
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

`/api/books/{_id}`に<b>GET</b>リクエストを送信すると、`title`、`_id`、および`comments`配列（コメントがなければ空の配列）を含む単一の本のオブジェクトが取得できます。本が見つからない場合は文字列`no book exists`を返してください。

```js
  try {
    let url = code + '/api/books';
    const noBookResponse = await fetch(url + '/5f665eb46e296f6b9b6a504d');
    if (!noBookResponse.ok) {
      throw Error(await noBookResponse.text());
    }
    let noBook = await noBookResponse.text();
    assert.isString(noBook);
    assert.equal(noBook, 'no book exists');
    const createResponse = await fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ title: 'Faux Book Alpha' })
    });
    if (!createResponse.ok) {
      throw Error(await createResponse.text());
    }
    let sampleBook = await createResponse.json();
    assert.isObject(sampleBook);
    let bookId = sampleBook._id;
    const queryResponse = await fetch(url + '/' + bookId);
    if (!queryResponse.ok) {
      throw Error(await queryResponse.text());
    }
    let bookQuery = await queryResponse.json();
    assert.isObject(bookQuery);
    assert.property(bookQuery, 'title');
    assert.equal(bookQuery.title, 'Faux Book Alpha');
    assert.property(bookQuery, 'comments');
    assert.isArray(bookQuery.comments);
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

`comment`に対して`/api/books/{_id}`をフォームボディデータとして含む<b>POST</b>リクエストを送信して、本にコメントを追加できます。返されるレスポンスは、以前のテストでの<b>GET</b> `/api/books/{_id}`リクエストに似た本のオブジェクトです。リクエストに`comment`が含まれていない場合は文字列`missing required field comment`を返してください。本が見つからない場合は文字列`no book exists`を返してください。

```js
  try {
    let url = code + '/api/books';
    const createResponse = await fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ title: 'Notable Book' })
    });
    if (!createResponse.ok) {
      throw Error(await createResponse.text());
    }
    let commentTarget = await createResponse.json();
    assert.isObject(commentTarget);
    let bookId = commentTarget._id;
    const comment1Response = await fetch(url + '/' + bookId, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ comment: 'This book is fab!' })
    });
    if (!comment1Response.ok) {
      throw Error(await comment1Response.text());
    }
    let bookCom1 = await comment1Response.json();
    const comment2Response = await fetch(url + '/' + bookId, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ comment: 'I did not care for it' })
    });
    if (!comment2Response.ok) {
      throw Error(await comment2Response.text());
    }
    let bookCom2 = await comment2Response.json();
    assert.isObject(bookCom2);
    assert.property(bookCom2, '_id');
    assert.property(bookCom2, 'title');
    assert.property(bookCom2, 'comments');
    assert.lengthOf(bookCom2.comments, 2);
    bookCom2.comments.forEach((comment) => {
      assert.isString(comment);
      assert.oneOf(comment, ['This book is fab!', 'I did not care for it']);
    });
    const commentErrResponse = await fetch(url + '/' + bookId, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({})
    });
    if (!commentErrResponse.ok) {
      throw Error(await commentErrResponse.text());
    }
    let commentErr = await commentErrResponse.text();
    assert.isString(commentErr);
    assert.equal(commentErr, 'missing required field comment');
    const failingCommentResponse = await fetch(url + '/5f665eb46e296f6b9b6a504d', {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ comment: 'Never Seen Comment' })
    });
    if (!failingCommentResponse.ok) {
      throw Error(await failingCommentResponse.text());
    }
    let failingComment = await failingCommentResponse.text();
    assert.isString(failingComment);
    assert.equal(failingComment, 'no book exists');
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

`/api/books/{_id}`に<b>DELETE</b>リクエストを送信して、コレクションから本を削除できます。成功した場合、返されるレスポンスは文字列`delete successful`です。本が見つからない場合は文字列`no book exists`を返してください。

```js
  try {
    let url = code + '/api/books';
    const createResponse = await fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ title: 'Deletable Book' })
    });
    if (!createResponse.ok) {
      throw Error(await createResponse.text());
    }
    let deleteTarget = await createResponse.json();
    assert.isObject(deleteTarget);
    let bookId = deleteTarget._id;
    const deleteResponse = await fetch(url + '/' + bookId, { method: 'DELETE' });
    if (!deleteResponse.ok) {
      throw Error(await deleteResponse.text());
    }
    let doDelete = await deleteResponse.text();
    assert.isString(doDelete);
    assert.equal(doDelete, 'delete successful');
    const failingDeleteResponse = await fetch(url + '/5f665eb46e296f6b9b6a504d', { method: 'DELETE' });
    if (!failingDeleteResponse.ok) {
      throw Error(await failingDeleteResponse.text());
    }
    let failingDelete = await failingDeleteResponse.text();
    assert.isString(failingDelete);
    assert.equal(failingDelete, 'no book exists');
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

`/api/books`に<b>DELETE</b>リクエストを送信して、データベース内のすべての本を削除できます。成功した場合、返されるレスポンスは文字列`complete delete successful`です。

```js
  try {
    const response = await fetch(code + '/api/books', { method: 'DELETE' });
    if (!response.ok) {
      throw Error(await response.text());
    }
    const deleteAll = await response.text();
    assert.isString(deleteAll);
    assert.equal(deleteAll, 'complete delete successful');
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

必要な10個の機能テストはすべて完了し、合格しています。

```js
  try {
    const response = await fetch(code + '/_api/get-tests');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const getTests = await response.json();
    assert.isArray(getTests);
    assert.isAtLeast(getTests.length, 10, 'At least 10 tests passed');
    getTests.forEach((test) => {
      assert.equal(test.state, 'passed', 'Test in Passed State');
      assert.isAtLeast(
        test.assertions.length,
        1,
        'At least one assertion per test'
      );
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```
