---
id: 6565a536ba1f9f25bd30e88b
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви реалізуєте алгоритм Луна. Цей алгоритм — формула для перевірки різних ідентифікаційних номерів.

Почніть із оголошення функції з назвою `main`, вона слугуватиме точкою входу програми. Використайте ключове слово `pass`, щоб уникнути помилки.

# --hints--

У вашому коді має бути `def main():`.

```js
({
    test: () => {
        assert(runPython(`
        import inspect
        inspect.isfunction(main)
        `))
    }
})
```

У вашому коді має бути `pass`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const pass = __helpers.python.getDef("\n"+transformedCode, "main");
        assert.include(pass?.function_body, "pass");
    }
})
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
