---
id: 656879c6f35a85a59c06b3a7
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Ви вже отримували доступ до елементів (символів) рядка, використовуючи оператор індексу `[]`. Також можна використовувати оператор індексу, щоб отримати діапазон символів у рядку за допомогою `string[start:stop:step]`:

```python
my_string = 'camperbot'
my_string[0:6] == 'camper' # True
my_string[0:6:3] == 'cp' # True
```

Де `start` — початковий індекс (включно), `stop` — кінцевий індекс (виключно), а `step` — кількість символів, які потрібно пропустити. Якщо не вказано, `step` за замовчуванням дорівнює 1.

Створіть змінну з ім’ям `card_number_reversed` і призначте їй значення перших 4 символів `card_number`.

# --hints--

У вас має бути `card_number_reversed = card_number[0:4]` всередині функції `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /card_number_reversed\s*=\s*card_number\s*\[\s*0\s*:\s*4\s*\]/);
    }
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
