---
id: 65687db6a9fcf3c548a6c1b7
title: Крок 28
challengeType: 20
dashedName: step-28
---

# --description--

Частина алгоритму полягає в тому, щоб подвоїти кожну другу цифру, починаючи з правого боку. Якщо результат подвоєння числа більший або дорівнює `10`, додайте дві цифри разом. Наприклад, якщо цифра `6`, подвойте її, щоб отримати `12`. Додайте `1` і `2` разом, щоб отримати `3`. Це можна зробити, використовуючи _ціле ділення_ для отримання першої цифри та оператор модуля (`%`) для отримання другої цифри:

```python
my_number = 12
first_digit = my_number // 10
second_digit = my_number % 10
```

Ціле ділення дає частку від ділення, округлену вниз до найближчого цілого числа.

У межах інструкції `if` призначте `number` результат `number // 10` (ціле ділення) плюс модуль `number` і `10`.

# --hints--

У вас має бути `number = number // 10 + number % 10` всередині інструкції `if`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getBlock("\n" + transformedCode, "if number >= 10");
        const { block_body } = verify_card_number;
        const allowedMatches = [
            /number *= *\(?\s*number *\/\/ *10\s*\)? *\+ *\(?\s*number *% *10/,
            /number *= *\(?\s*number *% *10\s*\)? *\+ *\(?\s*number *\/\/ *10/
        ];

        const someMatch = allowedMatches.some(allowedMatch => allowedMatch.test(block_body));
        assert.isTrue(someMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

--fcc-editable-region--
    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            print(number)
            
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
