---
id: 656877c202f44995ed952608
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Замініть інструкцію `pass` на змінну з іменем `card_number` та значенням `'4111-1111-4555-1142'`.

# --hints--

У вас має бути `card_number = '4111-1111-4555-1142'` всередині функції `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_number\s*=\s*('|")4111-1111-4555-1142\1/);
    }
})
```

У вашому коді не повинно бути `pass`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const no_comments = __helpers.python.removeComments(transformedCode);
        assert.notInclude(no_comments, "pass");
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    pass

--fcc-editable-region--
```
