---
id: 656877f71bba2b97acedf9af
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Python має вбудовані класи, які можуть допомогти нам із маніпуляціями рядками. Один із них — клас `str`. У нього є метод `maketrans`, який допомагає створити таблицю перекладу. Цю таблицю можна використовувати для заміни символів у рядку:

```python
str.maketrans({'t': 'c', 'l': 'b'})
```

Вищенаведене, коли викликається для рядка, замінить усі символи `t` на `c` і всі символи `l` на `b`.

Створіть змінну з назвою `card_translation` і призначте їй таблицю перекладу, щоб замінити всі символи `-` і ` ` на порожній рядок.

# --hints--

Ви повинні створити змінну `card_translation` всередині `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_translation\s*=/);
    }
})
```

Ви повинні призначити `card_translation` значення `str.maketrans({'-': '', ' ': ''})`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        const allowedMatches = [
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|")-\1\s*:\s*('|")\2\s*,\s*('|") \3\s*:\s*('|")\4\s*\}\s*\)/,
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|") \1\s*:\s*('|")\2\s*,\s*('|")-\3\s*:\s*('|")\4\s*\}\s*\)/,
        ];
        const anyMatch = allowedMatches.some((match) => match.test(function_body));
        assert.isTrue(anyMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    
--fcc-editable-region--
```
