---
id: 6569b831a07d7154c793301b
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Визначте функцію `verify_card_number` з параметром `card_number`, і використайте ключове слово `pass`, щоб функція нічого не робила.

# --hints--

У вас має бути функція з назвою `verify_card_number` з параметром `card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        assert.exists(verify_card_number);
        const { function_parameters } = verify_card_number;
        assert.equal(function_parameters.trim(), "card_number");
    }
})
```

У тілі функції `pass` слід використати ключове слово `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        assert.exists(verify_card_number);
        const { function_body } = verify_card_number;
        assert.match(function_body, / +pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--

def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

main()
```
