---
id: 67ef8f695db52583424bd118
title: Projetar uma Página de Avaliação de Filmes
challengeType: 25
dashedName: design-a-movie-review-page
demoType: onClick
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório. 

**Histórias de Usuário:**

1. Você deve ter apenas um elemento `main`.
2. Dentro do elemento `main`, você deve ter um elemento `h1` para o título do filme.
3. Abaixo do elemento `h1`, você deve ter um elemento `img` exibindo a capa do filme. Seu elemento `img` deve ter um texto `alt` descritivo que descreva a imagem. Você pode usar a seguinte imagem, se desejar: `https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png`.
4. Você deve ter um elemento `p` contendo uma breve descrição do filme.
5. Você deve ter outro elemento `p` para exibir a classificação do filme. Dentro dele, você deve ter estes itens na ordem listada:
   - Um elemento `strong` com o texto `Movie Rating:`.
   - Um elemento `span` com um atributo `aria-hidden` definido como `true` contendo uma representação visual da avaliação usando estrelas `⭐⭐⭐⭐⭐⭐⭐⭐⭐☆`.
   - Um valor numérico, representando a avaliação, entre parênteses (por exemplo, `9.2/10`) após o span.
6. Você deve ter um elemento `h2` com o texto `Membros do Elenco`.
7. Você deve ter um elemento `ul`.
8. Dentro do elemento `ul`, você deve ter múltiplos elementos `li`, cada um contendo um elemento `strong` para o nome do ator seguido pelo nome do personagem correspondente precedido pelo texto `as`. (ex.: `James Holloway as Ethan Carter`).

# --hints--

Você deve ter apenas um elemento `main`.

```js
assert.isNotNull(document.querySelector("main"));
```

Dentro do elemento `main`, você deve ter um elemento `h1` contendo o título do filme.

```js
const h1 = document.querySelector("main > h1");
assert.isAbove(h1?.innerText.trim().length, 0);
```

Abaixo do elemento `h1`, você deve ter um elemento `img` com um atributo `alt` descritivo.
  
```js
const img = document.querySelector("main > h1 + img");
assert.isNotNull(img);
assert.isNotEmpty(img.getAttribute("alt"));
```

Você deve ter um elemento `p` após a imagem que contenha a descrição do filme.

```js
const description = document.querySelector("main img + p");
assert.isNotNull(description);
assert.isNotEmpty(description.innerText.trim());
```

Você deve ter outro elemento `p` que contenha a avaliação.

```js
const paragraphs = document.querySelectorAll("main p");
assert.isAtLeast(paragraphs.length, 2);
```

Dentro do segundo elemento `p`, você deve ter um elemento `strong` com o texto `Classificação do Filme: `.</dfn>

```js
const strongEl = document.querySelector("main p:nth-of-type(2) strong");
assert.match(strongEl?.innerText.trim(), /^Movie Rating:?$/);
```

Deve haver um elemento `span` dentro do parágrafo de avaliação.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.isNotNull(spanEl);
```

O elemento `span` dentro do parágrafo de avaliação deve conter dez estrelas, preenchidas (⭐) ou vazias (☆), seguidas por um valor numérico entre parênteses após o `span`.

```js  
const spanEl = document.querySelector("main p:nth-of-type(2) span");
const fullRatingText = document.querySelector("main p:nth-of-type(2)")?.textContent.replace(/\s+/g, ' ').trim();

assert.match(spanEl?.innerText.trim(), /^[⭐☆]{10}$/);
assert.match(fullRatingText, /Movie Rating:?\s*[⭐☆]{10}\s*\(\s*\d+(?:\.\d+)?\s*\/\s*10\s*\)$/);
```

O elemento `span` dentro do parágrafo de avaliação deve ter um atributo `aria-hidden` definido como `true`.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.strictEqual(spanEl?.getAttribute("aria-hidden"), "true");
```

Você deve ter um elemento `h2`.

```js
assert.exists(document.querySelector("h2"));
```

Seu `h2` deve vir após seu segundo elemento `p`.

```js
const allParagraphs = document.querySelectorAll("p");
assert.isAtLeast(allParagraphs.length, 2);

const h2 = document.querySelector("h2");
const main = document.querySelector("main");

assert.exists(h2);
assert.exists(main);

const elements = Array.from(main.querySelectorAll("*"));
const lastPIndex = elements.indexOf(allParagraphs[1]);
const h2Index = elements.indexOf(h2);

assert.isAbove(h2Index, lastPIndex);
```

Seu `h2` deve conter o texto `Membros do Elenco`.

```js
const h2 = document.querySelector("h2");
assert.equal(h2?.textContent.trim(), "Cast Members");
```

Você deve ter um elemento `ul` após o `h2`.

```js
const ul = document.querySelector("main h2 + ul");
assert.isNotNull(ul);
```

O `ul` deve conter múltiplos elementos `li`.

```js
const liItems = document.querySelectorAll("main ul li");
assert.isAbove(liItems.length, 1);
```

Cada `li` deve conter um elemento `strong` para o nome do ator seguido pelo nome do personagem correspondente precedido pelo texto `as`. (ex.: `James Holloway as Ethan Carter`).

```js
const liItems = document.querySelectorAll("main ul li");  
assert.isNotEmpty(liItems);
liItems.forEach(li => {
  const strong = li.querySelector("strong");
  assert.isNotNull(strong);
  const fullText = li.textContent.trim();
  const actorName = strong.textContent.trim();
  assert.isTrue(fullText.startsWith(actorName));
  const characterPart = fullText.slice(actorName.length).trim();
  assert.match(characterPart, /^as\s.+/);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
 <html lang="en">
 <head>
     <meta charset="UTF-8">
     <meta name="viewport" content="width=device-width, initial-scale=1.0">
     <title>Movie Review</title>
 </head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Movie Review</title>
  </head>
  <body>
    <main class="review">
      <h1>Movie Review: Rise Beyond</h1>
      <img
        src="https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png"
        alt="The Rise Beyond movie poster, starring Ethan Carter. A portrait of a young boy standing on a stage with his back toward the audience, holding a guitar at his side, staring at a backdrop of a cityscape with the sun rising behind its tall buildings."
      />
      <p>
        Rise Beyond is an uplifting drama that tells the story of Ethan Carter,
        a young man from a small town who dreams of becoming a world-class
        musician. Facing financial struggles and self-doubt, Ethan embarks on a
        journey of resilience, guided by his mentor Professor Adams and
        childhood friend Lena Mitchell. Through setbacks, failures, and moments
        of triumph, Rise Beyond reminds us that our greatest potential lies
        beyond our fears.
      </p>
      <p>
        <strong>Movie Rating:</strong>
        <span aria-hidden="true">⭐⭐⭐⭐⭐⭐⭐⭐⭐☆</span> (9.2/10)
      </p>

      <h2>Cast Members</h2>
      <ul>
        <li><strong>James Holloway</strong> as Ethan Carter</li>
        <li><strong>Olivia Sterling</strong> as Lena Mitchell</li>
        <li><strong>William Lancaster</strong> as Professor Adams</li>
        <li>
          <strong>Maria Collins</strong> as Evelyn Carter (Ethan’s Mother)
        </li>
      </ul>
    </main>
  </body>
</html>
```
