---
id: 683ec95f2f8781355556ff81
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

Unahitaji kujaribu mambo ili uelewe mtiririko wa algoriti katika hatua hii ya mwanzo.

Ili kufanya hivyo, unahitaji kwanza kutoka kwenye mzunguko. Hii ni kwa sababu utekelezaji wa sasa utaruhusu mzunguko mmoja tu, hivyo kama sharti halitimizwi, kutakuwa na mzunguko usioisha.

Mara tu baada ya `if`statement, tumia neno kuu la `break` kutoka kwenye mzunguko wa `while`. Kisha, baada ya mzunguko wa `while`, rudisha orodha tupu kuashiria kuwa thamani haikupatikana.

# --hints--

Unapaswa kutumia neno kuu la `break` kutoka kwenye mzunguko wa `while` baada ya tamko la `if`.

```js
({
    test: () => assert(runPython(
      `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().has_stmt("break")`)
    )
})
```

Unapaswa kurudisha orodha tupu baada ya mzunguko wako wa `while`.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").has_return("[]")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--
```
