---
id: 683ec95f2f8781355556ff7e
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Sasa umehesabu kielezo cha `mid`, unahitaji sasa kupata thamani kutoka `search_list` katika kielezo hicho. Thamani hii ndiyo utakayolinganishwa na kigezo chako cha `value`, ambacho ni lengo unalotafuta.

Tengeneza tofauti ya `value_at_middle` na weka thamani yake kuwa kipengele kutoka `search_list` kilichoko katika kielezo cha `mid`.

# --hints--

Unapaswa kutengeneza tofauti ya `value_at_middle` ambayo imewekwa thamani ya `search_list[mid]`.

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid]")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
--fcc-editable-region--
        
--fcc-editable-region--
```
