---
id: 683ec95f2f8781355556ff7d
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Ili kuanza na mzunguko, unahitaji kupata kielezo cha katikati cha nafasi ya sasa ya utafutaji.

Unaweza kupata kielezo cha katikati kwa kuhesabu wastani wa vigezo vya `low` na `high`. Itabidi utumie mgawanyiko wa floor ili kupata wastani baada ya kuongeza thamani hizo mbili, hivyo jibu litakuwa daima nambari kamili iliyozungushwa chini, kwani vielezo lazima viwe nambari kamili.

Ondoa neno la `pass` kutoka kwenye mzunguko wa `while` na unda kigezo cha `mid` kilichowekwa kwa wastani wa `low` na `high`.

# --hints--

Unapaswa kuondoa neno la `pass` kutoka kwenye mzunguko wa while na kuunda kigezo cha `mid` kilichowekwa kwa `(low + high) // 2`.

```js
({
    test: ()=>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2") or _Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (high + low) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

--fcc-editable-region--
    while low <= high:
        pass
--fcc-editable-region--
```
