---
id: 683ec95f2f8781355556ff7b
title: Hatua ya 3
challengeType: 20
dashedName: step-3
---

# --description--

Katika utafutaji wa binary, kuna kielezo cha chini kabisa kinachowezekana na kielezo cha juu kabisa kinachowezekana.

Kielezo cha chini kabisa kinachowezekana kinakilisha mpaka wa kushoto kabisa katika sehemu ya sasa ya maadili yanayotafutwa, katika kesi hii, `search_list`. Kielezo cha juu kabisa kinachowezekana, kwa upande mwingine, kinakilisha mpaka wa kulia kabisa katika maadili yanayotafutwa.

Ili kuzingatia hayo mawili, lazima uzingatie orodha nzima inayotafutwa. Hivyo, ndani ya kitendakazi cha `binary_search`, fafanua kigezo cha `low` chenye thamani ya `0`, na kigezo cha `high` chenye thamani inayozingatia kielezo cha mwisho katika orodha inayotafutwa.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `low` katika kitendakazi chako cha `binary_search`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("low")`)) 
})
```

Kigezo chako cha `low` kinapaswa kuwekwa kuwa `0`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_variable("low").is_equivalent("low = 0")`
      )
    )
})
```

Unapaswa kuwa na kigezo kinachoitwa `high` katika kitendakazi chako cha `binary_search`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("high")`)) 
})
```

Kigezo chako cha `high` kinapaswa kuwekwa kuwa `len(search_list) - 1`.

```js
({
  test: () => assert(
    runPython(
    `_Node(_code).find_function("binary_search").find_variable("high").is_equivalent("high = len(search_list) - 1")`
    ))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    path_to_target = []
    
--fcc-editable-region--
```
