---
id: 683ec95f2f8781355556ff7c
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Sasa baada ya kufafanua mipaka ya utafutaji wako, ni wakati wa kuunda mzunguko utakaotekeleza algoriti ya utafutaji wa binary.

Utafutaji wa binary hufanya kazi kwa kupunguza mara kwa mara nafasi ya utafutaji. Mchakato huu unaendelea mradi tu kuna safu halali ya vipengele vya kukagua.

Unaweza kudhihirisha hili kwa mzunguko wa `while` ambao utaendelea mradi tu kiashiria chako cha `low` kiko chini au sawa na kiashiria cha `high`.

Ikiwa `low` itakuwa kubwa kuliko `high`, inamaanisha nafasi ya utafutaji imeisha, na `value` haipo kwenye orodha.

Ndani ya kitendakazi cha `binary_search`, unda mzunguko wa `while` wenye sharti la kukagua kama `low` iko chini au sawa na high. Ndani ya mzunguko huo, ongeza neno kuu la `pass` kama kiwakilishi kwa sasa.

# --hints--

Unapaswa kuunda mzunguko wa `while` unaoendelea mradi `low` iko chini au sawa na `high`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_conditions()[0]
assert _cond.is_equivalent("low <= high") or _cond.is_equivalent("high >= low")`
) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
--fcc-editable-region--
    
--fcc-editable-region--
```
