---
id: 683ec95f2f8781355556ff7a
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Algoriti ya utafutaji wa binary unayojenga haitarudisha tu `True` au `False` kuonyesha kama thamani ilipatikana au la. Pia itarudisha njia inayochukua kupata thamani hiyo. Na kwa ajili ya hilo, utakuwa na tofauti ya `path_to_target` kufuatilia njia hiyo.

Ndani ya kitendakazi chako cha `binary_search`, badilisha neno la `pass` na tofauti ya `path_to_target` na uanze na orodha tupu.

# --hints--

Unapaswa kuunda tofauti ya `path_to_target` iliyowekwa kuwa orodha tupu ndani ya kitendakazi chako cha `binary_search`. Usisahau kuondoa neno la `pass`.

```js
({ 
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_body().is_equivalent("path_to_target = []") and not _Node(_code).find_function("binary_search").has_pass()`
      )
    ) 
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
