---
id: 683ec95f2f8781355556ff83
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Unaweza kuona kwamba wito la kwanza hurudisha `[3]`. Hii ni kwa sababu `3` ni thamani ya katikati katika `[1, 2, 3, 4, 5]`.

Kwa upande mwingine, wito la pili hurudisha orodha tupu. Hii hutokea kwa sababu, baada ya kukagua thamani ya katikati, mzunguko kwa sasa hauna mantiki ya kupunguza eneo la utafutaji kwa kusahihisha vigezo vya `low` au `high`. Inafanya tu ukaguzi mmoja wa katikati kisha inamaliza bila utafutaji zaidi.

Ili kuruhusu utafutaji wa binary kuendelea kupunguza utafutaji wake, ongeza kipengele cha `elif` kinachokagua kama `value` ni kubwa kuliko `value_at_middle`. Ongeza neno kuu la `pass` ndani ya `elif` kwa sasa.

# --hints--

Unapaswa kuongeza kipengele cha `elif` wakati `value` ni kubwa kuliko `value_at_middle`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("value > value_at_middle") or _cond.is_equivalent("value_at_middle < value")
`) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--          
        break
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
```
