---
id: 683ec39565dbc7e3375263b1
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika funzo lililopita, ulijifunza kuhusu utafutaji wa binary na jinsi unavyotofautiana na utafutaji wa mstari.

Katika warsha hii, utaweka hayo katika vitendo kwa kujenga algoriti ya utafutaji wa binary.

Kuanza, tengeneza kitendakazi cha `binary_search` chenye vigezo vya `search_list` na `value`.

`search_list` ni orodha ya vipengele ambavyo kitendakazi kitaangalia ndani yake, na `value` ni lengo ambalo kitendakazi kitajaribu kulipata ndani ya `search_list`.

Ongeza neno kuu la `pass` ndani ya kitendakazi kwa sasa ili jaribio lipite.

# --hints--

Unapaswa kutengeneza kitendakazi cha `binary_search`.

```js
({ 
  test: () => assert(runPython(`_Node(_code).has_function("binary_search")`)) 
})
```

Kitendakazi chako cha `binary_search` kinapaswa kuwa na vigezo vya `search_list` na `value`.

```js
({ 
  test: () => assert(
    runPython(
      `_Node(_code).find_function("binary_search").has_args("search_list, value") or _Node(_code).find_function("binary_search").has_args("value, search_list")`
    )
  ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
