---
id: 683ec95f2f8781355556ff7f
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Sasa, jumuisha `value_at_middle` katika orodha ya `path_to_target` kufuatilia hatua zilizochukuliwa wakati wa utafutaji, bila kujali kama ni lengwa.

# --hints--

Unapaswa kutumia njia ya `append()` kuongeza `value_at_middle` kwenye orodha ya `path_to_target`.

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid] \\npath_to_target.append(value_at_middle)")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
--fcc-editable-region--
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        
--fcc-editable-region--
```
