---
id: 683ec95f2f8781355556ff82
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Kumbuka kuwa kitendakazi kinapokea orodha iitwayo `search_list` na `value` unayotafuta.

Sasa, fanya wito wa kitendakazi kwa `binary_search([1, 2, 3, 4, 5], 3)` na `binary_search([1, 2, 3, 4, 5, 9], 4)` na chapisha wito hizo mara moja.

# --hints--

Unapaswa kufanya wito wa kitendakazi kwa `binary_search([1, 2, 3, 4, 5], 3)` na kuichapisha mara moja.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5], 3))")
    `))
})
```

Unapaswa kufanya wito wa kitendakazi kwa `binary_search([1, 2, 3, 4, 5, 9], 4)` na kuichapisha mara moja.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5, 9], 4))")
    `))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
        
        if value == value_at_middle:
            return path_to_target
        break
    
    return []

--fcc-editable-region--

--fcc-editable-region--
```
