---
id: 684169b00a073c9b9f0a703c
title: Hatua 14
challengeType: 20
dashedName: step-14
---

# --description--

Je, umeona wito wa pili wa kitendakazi sasa una thamani `3, 5, 4` kwenye orodha? Hiyo inaonyesha kuwa utafutaji wa binary unafanya kazi kama ilivyokusudiwa.

Hivi ndivyo ilivyotokea: Algoriti ilikagua kwanza `3` kama katikati ya orodha ya awali. Kwa kuwa `4` ni kubwa kuliko `3`, utafutaji uligeuka upande wa kulia. Kisha ukachunguza `5` kama katikati mpya. Kwa sababu `4` ni ndogo kuliko `5`, utafutaji uligeuka upande wa kushoto, hatimaye ukatambua `4` kama katikati ya safu ya mwisho.

Ili kujaribu kitendakazi tena, kiite kwa `[1, 3, 5, 9, 14, 22], 10` na chapisha wito mara moja. Hii ni hali ambapo thamani haitapatikana.

# --hints--

Unapaswa kuita kitendakazi kwa `binary_search([1, 3, 5, 9, 14, 22], 10)` na kukichapisha mara moja.

```js
({
    test: () => assert(runPython(
        `_Node(_code).has_call("print(binary_search([1, 3, 5, 9, 14, 22], 10))")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1

    return []
    
print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
--fcc-editable-region--

--fcc-editable-region--
```
