---
id: 662a6bc12cde72c32fb526f0
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

<dfn>介面</dfn>就像是類別的藍圖。介面包含一組類別應該實作的方法和屬性。  

從宣告一個名為 `Equation` 的空類別開始這個專案。你將使用這個類別來定義一個介面，一個泛型方程式的藍圖。

# --hints--

你應該定義一個名為 `Equation` 的新類別。

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
