---
id: 662bde88dc84f1e249801b1a
title: 步驟6
challengeType: 20
dashedName: step-6
---

# --description--

`ABC` 代表 *Abstract Base Classes*。 `ABC` 類使你能夠將常規類轉換爲抽象類，抽象類可作爲具體類的藍圖。

修改 `import` 語句，只從 `abc` 模塊導入 `ABC` 類。 你可以使用以下導入語句 `from y import x` 從模塊 `y` 導入特定對象 `x`。

然後，通過繼承 `ABC` 類將您的 `Equation` 類變成抽象類。

# --hints--

你應該從 `abc` 模塊中導入 `ABC`。

```js
({ test: () => assert(runPython(`_Node(_code).has_import("from abc import ABC")`)) })
```

你的 `Equation` 類應該繼承 `ABC`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").inherits_from("ABC")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import abc


class Equation:
    def solve(self):
        pass

    def analyze(self):
        pass

class LinearEquation(Equation):
    pass

eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
