---
id: 662bd456896f16d9bd03f1a6
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

在 `Equation` 類別中，定義兩個名為 `solve` 和 `analyze` 的新實體方法。

# --hints--

你應該在 `Equation` 類別中定義一個名為 `solve` 的方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("solve")`)) })
```

你的 `solve` 方法應該接受一個參數，`self`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_args("self")`)) })
```

你應該在 `Equation` 類別中定義一個名為 `analyze` 的方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("analyze")`)) })
```

你的 `analyze` 方法應該接受一個參數，`self`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Equation:
    pass
--fcc-editable-region--
```
