---
id: 662bd8260da84bdd5feae419
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

你想讓 `LinearEquation` 類別實作而非單純繼承 `Equation` 類別內定義的所有方法，該類別應該作為一個介面。

目前，`Equation` 類別只是 `LinearEquation` 的父類別。在接下來的步驟中你將學習如何將它轉換成正式的介面。

目前，創建一個 `Equation` 的實體並將它指定給變數 `eq`，以及創建一個 `LinearEquation` 的實體並將它指定給變數 `lin_eq`。

# --hints--

你應該宣告一個變數 `eq` 並指定它為一個 `Equation` 的實體。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("eq = Equation()")`)) })
```

你應該宣告一個變數 `lin_eq` 並指定它為一個 `LinearEquation` 的實體。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation()")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
--fcc-editable-region--

--fcc-editable-region--
```
