---
id: 662bd552e1c1d2db1b88ba47
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

現在，定義另一個類名爲 `LinearEquation` ，並使它繼承 `Equation`。 你將使用此類表示線性方程。

# --hints--

你應該定義一個名爲 `LinearEquation` 的類。

```js
({ test: () => assert(runPython(`_Node(_code).has_class("LinearEquation")`)) })
```

`LinearEquation` 類應該繼承 `Equation` 類。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").inherits_from("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass

    def analyze(self):
        pass
--fcc-editable-region--

--fcc-editable-region--
```
