---
id: 662bde88dc84f1e249801b1a
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

`ABC` 代表 *抽象的基礎類別*。`ABC` 類別啟用你將一般類別轉換成抽象類別，抽象類別是作為具體類別藍圖的類別。

修改你的 `import` 述語以僅從 `abc` 模組匯入 `ABC` 類別。你可以依照匯入構件 `from y import x` 從模組 `y` 匯入特定物件 `x`。

然後，透過讓你的 `Equation` 類別繼承自 `ABC`，將它轉換成抽象的類別。

# --hints--

你應該從 `abc` 模組匯入 `ABC`。

```js
({ test: () => assert(runPython(`_Node(_code).has_import("from abc import ABC")`)) })
```

你的 `Equation` 類別應該繼承自 `ABC`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").inherits_from("ABC")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import abc


class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass
    
eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
