---
id: 662bd552e1c1d2db1b88ba47
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

現在，定義另一個名為 `LinearEquation` 的類別，並讓它繼承自 `Equation`。你將使用這個類別來表述線性方程式。

# --hints--

你應該定義一個名為 `LinearEquation` 的類別。

```js
({ test: () => assert(runPython(`_Node(_code).has_class("LinearEquation")`)) })
```

你的 `LinearEquation` 類別應該繼承自 `Equation` 類別。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").inherits_from("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
--fcc-editable-region--

--fcc-editable-region--
```
