---
id: 662bdd364bf2cde1487922a9
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

與其他程式設計語言不同，Python 並未在其核心語言中實作介面，但 Python 標準程式庫允許你以簡單的方式定義介面。

在這個專案中，你將使用來自 `abc` 模組的工具。因此，請在你的程式碼中匯入此模組。

# --hints--

你應該匯入 `abc` 模組。

```js
({ test: () => assert(runPython(`_Node(_code).has_import("import abc")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
    

eq = Equation()
lin_eq = LinearEquation()
```
