---
id: 662fc3eba556a6bf800d48c1
title: 步驟14
challengeType: 20
dashedName: step-14
---

# --description--

`isinstance()` 內置函數接收兩個參數，並返回一個布爾值，表示作爲第一個參數傳遞的對象是否是作爲第二個參數傳遞的類的實例。

```py
isinstance(7, int) # True
```

你需要檢查的另一件事是每個參數都是一個數字。 在第一個 `if` 之後，創建一個 `for` 循環，迭代 `args` 並檢查當前迭代的參數是不是 `int` 或 `float` 的實例。 使用 `isinstance()` 函數並傳遞一個包含 `int` 和 `float` 的元組作爲第二個參數。

如果參數不是數字，則引發 `TypeError`，提示 `"Coefficients must be of type 'int' or 'float'"`。

# --hints--

你應該在 `if` 語句後創建一個 `for` 循環來迭代 `args`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_iter().is_equivalent("args")`)) })
```

你應該創建一個 `if` 語句，檢查當前係數是不是 `for` 循環中的 `int` 或 `float` 的實例。

```js
({ test: () => assert(runPython(`
var = str(_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_vars())
cond1 = f'not isinstance({var}, (int, float))'
cond2 = f'not isinstance({var}, (float, int))'
if_stmt = _Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_conditions()[0]
if_stmt.is_equivalent(cond1) or if_stmt.is_equivalent(cond2)
`)) })
```

你應該使用提供的字符串在 `if` 語句中引發 `TypeError`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_bodies()[0].has_stmt("raise TypeError(\\"Coefficients must be of type 'int' or 'float'\\")")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    degree: int

    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    @abstractmethod
    def solve(self):
        pass

    @abstractmethod
    def analyze(self):
        pass

class LinearEquation(Equation):
    degree = 1

    def solve(self):
        pass

    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
```
