---
id: 662fa2e2cf27c09f21f4f5d0
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

在 Python 中，資料型別是在執行期（當程式碼被執行時）被識別的。因此，你在宣告變數時不需要指定資料型別。儘管如此，你可以使用 `variable: <data type> = value` 或僅 `variable: <data type>` 來註解變數，以說明它將會儲存特定的資料型別。請注意，Python 直譯器不會強制執行用於註解變數的型別，通常你需要外部工具來完成此工作。

在 `Equation` 類別內定義一個類別屬性 `degree`。不要為它指定值。改用 `int` 的型別註解來表示它將在具體的類別中儲存一個整數。

稍後，你會將這個類別屬性用作驗證傳遞以實例化方程式物件的引數的部分處理程序。

# --hints--

你應該在 `Equation` 類別中定義名為 `degree` 的類別屬性，並以 `int` 註解它。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_variable("degree").is_equivalent("degree: int")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    def __init__(self):
        pass
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass
--fcc-editable-region--
lin_eq = LinearEquation()
```
