---
id: 662f96576ef178927de87975
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

為了被識別為抽象的行為，該行為應該使用 `@abstractmethod` 裝飾器裝飾。

<dfn>裝飾器</dfn> 用於 Python 中修改函式（程式）的行為。以下是如何使用名為 `spam` 的裝飾器的範例：

```py
@spam
def foo():
    pass
```

修改你的匯入述語以匯入 `abstractmethod` 裝飾器，並裝飾 `Equation` 類別的 `solve` 和 `analyze` 兩個方法。這將引發兩個例外。

一旦繼承自 `ABC` 的類別擁有抽象的方法，該類別就無法再被具現化。因此，刪除 `Equation` 實體以消除錯誤。

另一個錯誤發生是因為 `LinearEquation` 類別必須實作介面中定義的所有抽象的行為。請確定也在 `LinearEquation` 類別內定義它們。你不得在具體的類別中使用 `abstractmethod` 裝飾器。

# --hints--

你應該從 `abc` 模組匯入 `abstractmethod`。

```js
({ test: () => assert(runPython(`
_Node(_code).has_import("from abc import ABC, abstractmethod") or \\
_Node(_code).has_import("from abc import abstractmethod, ABC") or \\
(_Node(_code).has_import("from abc import abstractmethod") and _Node(_code).has_import("from abc import ABC"))
`)) })
```

你應該在 `Equation` 類別中的 `solve` 方法上使用 `@abstractmethod` 裝飾。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_decorators("abstractmethod")`)) })
```

你應該在 `Equation` 類別中的 `analyze` 方法上使用 `@abstractmethod` 裝飾。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_decorators("abstractmethod")`)) })
```

你應該在 `LinearEquation` 類別中定義一個名為 `solve` 的方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("solve")`)) })
```

你的 `solve` 方法應該接受一個參數，`self`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_args("self")`)) })
```

你應該在 `LinearEquation` 類別中定義一個名為 `analyze` 的方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("analyze")`)) })
```

你的 `solve` 方法應該接受一個參數，`self`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC


class Equation(ABC):
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass

eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
