---
id: 662f6d7c92381a3049e4c987
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

介面不必只定義抽象的行為，也可以實作行為，讓具體的類別繼承。

在處理 `Equation` 類別中 `solve` 和 `analyze` 的實作之前，先定義一個 `__init__` 方法。不要在它上面使用任何裝飾器。

# --hints--

你應該在你的 `Equation` 類別中定義一個 `__init__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__init__")`)) })
```

你的 `__init__` 方法應該接受一個參數，`self`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC, abstractmethod


class Equation(ABC):
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
--fcc-editable-region--
```
