---
id: 662fbcef5f05e1b84f541a0c
title: 步驟 13
challengeType: 20
dashedName: step-13
---

# --description--

每個方程式物件將會具現化，並傳入與方程式係數數量相同的引數，從 \\( x \\) 的 n 次方開始，直到零次方，包括可能值為 `0` 的係數。

例如，`LinearEquation(4, 5)` 將表述方程式 \\( 4x + 5 = 0 \\)，其中 `4` 是一階（此處最高階）的係數，`5` 是零階的係數。

你需要檢查傳遞給方程式物件的引數數量是否正確。

在 `__init__` 方法內，創建一個 `if` 述語來檢查 `args` 的長度是否與方程式應有的係數數量（`degree + 1`）不同。如果不同，則引發 `TypeError`，並使用以下字串提供訂製訊息：`f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"`。

然後，透過傳入 `2` 和 `3` 來實例化 `lin_eq`，以修正錯誤。

# --hints--

你應該創建一個 `if` 述語，用來檢查用於實例化方程式的係數數量是否不同於 `degree + 1`。

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_conditions()[0]
cond.is_equivalent("(self.degree + 1) != len(args)") or cond.is_equivalent("len(args) != (self.degree + 1)")
`)) })
```

你應該在新的 `if` 述語中引發 `TypeError`，並使用提供的字串來傳回訂製的錯誤訊息。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError(f"\\'{self.__class__.__name__}\\' object takes {self.degree + 1} positional arguments but {len(args)} were given")')
`)) })
```

你應該傳入 `2` 和 `3` 來實例化 `lin_eq`。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation(2, 3)")`)) })
```


# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        pass        
    
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation()
--fcc-editable-region--
```
