---
id: 6639fdcc701833a54c364211
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

在驗證係數後，你需要將它們儲存在實體屬性中。使用字典推導式創建一個字典，其中鍵是係數的次數，對應的值是該係數，並將其指定給名為 `coefficients` 的屬性。

例如，使用 `2` 和 `4` 具現化的 `LinearEquation` 物件應該有以下的 `coefficients` 屬性：`{1: 2, 0: 4}`，因為 `2` 對應到 `x` 的一次方，`4` 對應到 `x` 的零次方。

按照 `args` 中的相同順序，在你的新字典中創建鍵值對。

# --hints--

你應該在你的 `__init__` 方法中宣告一個名為 `coefficients` 的屬性。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_variable("self.coefficients")`)) })
```

你應該使用字典推導式來儲存你的係數。

```js
({ test: () => runPython(`
import ast
node = _Node(_code).find_class("Equation").find_function("__init__").find_variable("self.coefficients")
assert isinstance(node.tree.value, ast.DictComp)
`) })
```

你的 `coefficients` 屬性應該是一個字典，包含以 degree-coefficient 形式的鍵值對組。請記得遵循 `args` 中儲存係數的相同順序。

```js
({ test: () => runPython(`
actual1 = list(LinearEquation(1, 6).coefficients.items())
expected1 = list({1: 1, 0: 6}.items())
actual2 = list(LinearEquation(-3.5, 0).coefficients.items())
expected2 = list({1: -3.5, 0: 0}.items())
assert actual1 == expected1
assert actual2 == expected2
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
--fcc-editable-region--
        
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass

    
lin_eq = LinearEquation(2, 3)
```
