---
id: 663a22ba7420c4d2f7fd2aec
title: 步驟 25
challengeType: 20
dashedName: step-25
---

# --description--

現在是實現 `solve` 方法的時候了。 給定形式爲 \\( ax + b = 0 \\) 的線性方程，其解爲 \\(x = -\frac{b}{a}\\)。

將存儲在 `coefficients` 屬性中的係數解包到變量 `a` 和 `b` 中。 請注意，你需要使用 `.values()` 方法。

然後，聲明一個變量 `x`，將方程的解賦給它並從 `solve` 方法返回它。

# --hints--

你應該將存儲在 `coefficients` 屬性內的值解包到變量 `a` 和 `b` 中。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_stmt("a, b = self.coefficients.values()")`)) })
```

你應該聲明一個名爲 `x` 的變量，併爲其分配線性方程的解。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_stmt("x = -b/a")`)) })
```

你應該從 `solve` 方法中返回 `x`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_return("x")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int

    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        

    @abstractmethod
    def solve(self):
        pass

    @abstractmethod
    def analyze(self):
        pass

class LinearEquation(Equation):
    degree = 1
--fcc-editable-region--
    def solve(self):
        pass
--fcc-editable-region--

    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
print(lin_eq)

```
