---
id: 663b7fefd437bd984e091cbf
title: 第 29 步
challengeType: 20
dashedName: step-29
---

# --description--

接著，創建一個名為 `QuadraticEquation` 的新類別，並讓它繼承自 `Equation`。你將使用這個新類別來表述二次方程式，二次方程式是具有形式 $ax^2 + bx + c = 0$ 的二次方程。 

在你的新類別中，定義一個值為 `2` 的 `degree` 類別屬性，這是二次方程式的次數。也定義 `solve` 和 `analyze` 方法。你將在接下來的步驟中負責實作。

# --hints--

你應該創建一個名為 `QuadraticEquation` 的新類別，並讓它繼承自 `Equation` 類別。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").inherits_from("Equation")`)) })
```

你應該在 `QuadraticEquation` 類別中定義一個 `solve` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").has_function("solve")`)) })
```

你的 `solve` 方法應該接受一個單一參數，`self`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("solve").has_args("self")`)) })
```

你應該在 `QuadraticEquation` 類別中定義一個 `analyze` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").has_function("analyze")`)) })
```

你的 `analyze` 方法應該接受一個單一的參數，`self`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("analyze").has_args("self")`)) })
```

你應該在 `QuadraticEquation` 類別中定義一個 `degree` 類別屬性，並將其指定為值 `2`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_variable("degree").is_equivalent("degree = 2")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}
--fcc-editable-region--

--fcc-editable-region--
lin_eq = LinearEquation(2, 3)
print(lin_eq)

```
