---
id: 663b83a28943e6aa6275a514
title: 第 19 步
challengeType: 20
dashedName: step-19
---

# --description--

仍在 `Equation` 類別中，定義一個 `__str__` 方法，為你將要創建的方程式物件提供適當的字串表現。

目前，在 `__str__` 方法中，宣告一個變數 `terms` 並指定為空列表。你將使用這個變數來儲存方程式中的每一項（係數乘以 \\( x^n \\)）。

然後，宣告一個變數 `equation_string`，將 `terms` 列表中元素以空格串接的結果指定給它。最後，傳回 `equation_string`。

# --hints--

你應該在 `Equation` 類別中定義一個 `__str__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__str__")`)) })
```

你的 `__str__` 方法應該接受一個參數，`self`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_args("self")`)) })
```

你應該在 `__str__` 方法中宣告一個變數 `terms` 並指定一個空的列表。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("terms = []")`)) })
```

你應該在 `__str__` 方法中宣告一個變數 `equation_string`，並將以空格連接 `terms` 中元素的結果指定給它。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("equation_string = ' '.join(terms)")`)) })
```

你應該從你的 `__str__` 方法傳回 `equation_string`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_return("equation_string")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    
--fcc-editable-region--    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
