---
id: 664f0389424a6f7aa15fd3e5
title: 第 36 步
challengeType: 20
dashedName: step-36
---

# --description--

另一種環視斷言是前瞻。正向和負向前瞻分別以 `(?=...)` 和 `(?!...)` 表示。它們用於匹配後面跟隨特定字元序列的樣式，但該序列不會被消耗：

```py
spam = 'black back bat'
re.sub('a(?=t)', 'o', spam) == 'black back bot' # True
re.sub('a(?!t)', 'o', spam) == 'block bock bat' # True
```

在上述範例中，樣式 `a(?=t)` 包含正向前瞻，用來匹配僅在後面跟著 `t` 時的 `a` 字元。在範例的最後一行，樣式 `a(?!t)` 包含負向前瞻，用來匹配僅在後面**不**跟著 `t` 時的 `a` 字元。同樣地，在這兩種情況中，前瞻中包含的字元不會被消耗。

為你的 regex 樣式添加正向前瞻，使字元 `1` 僅在後面跟著字元 `x` 時才被取代。

# --hints--

你應該使用正向前瞻來修改你的 regex 樣式，以便只有在 `1` 後面跟著 `x` 時才取代該字元。不要從你的樣式中移除負向後顧。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_return("re.sub(r'(?<!\\d)1(?=x)', '', equation_string.strip('+'))")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
import re


class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')
            else:
                terms.append(f"{coefficient:+}x**{n}")
        equation_string = ' '.join(terms) + ' = 0'
--fcc-editable-region--
        return re.sub(r'(?<!\d)1', '', equation_string.strip('+'))        
--fcc-editable-region--
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}

class QuadraticEquation(Equation):
    degree = 2

    def __init__(self, *args):
        super().__init__(*args)
        a, b, c = self.coefficients.values()
        self.delta = b**2 - 4 * a * c
   
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
quadr_eq = QuadraticEquation(11, -1, 1)
print(quadr_eq)

```
