---
id: 66562f71937f877c66123bbe
title: 步驟 49
challengeType: 20
dashedName: step-49
---

# --description--

輸出頂部要顯示的第一件事是方程式型別。為你的 `Equation` 類別添加一個名為 `type` 的類別屬性，並用 `str` 註解它。

然後，為 `__init_subclass__` 方法添加另一個 `if` 述語，以檢查繼承自 `Equation` 的類別是否具有 `type` 屬性。使用現有 `if` 述語的相同格式並進行適當修改。

最後，將新的類別屬性添加到 `LinearEquation` 類別和 `QuadraticEquation` 類別。分別為它指定字串 `'Linear Equation'` 和字串 `'Quadratic Equation'`。

# --hints--

你應該在 `Equation` 類別中定義一個名為 `type` 的類別變數，並用 `str` 註解它。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_stmt("type: str")`)) })
```

你應該在 `__init_subclass__` 方法中創建一個 `if` 述語，檢查 `cls` 是否沒有屬性 `type`，並使用提供的字串引發 `AttributeError`。

```js
({ test: () => assert(runPython(`
if_str = """
if not hasattr(cls, 'type'):
  raise AttributeError(
    f\\"Cannot create '{cls.__name__}' class: missing required attribute 'type'\\"
  )
"""
_Node(_code).find_class("Equation").find_function("__init_subclass__").has_stmt(if_str)
`)) })
```

`LinearEquation` 類別的 `type` 屬性應該具有值 `'Linear Equation'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_stmt("type = 'Linear Equation'")`)) })
```

`QuadraticEquation` 類別的 `type` 屬性應該具有值 `'Quadratic Equation'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").has_stmt("type = 'Quadratic Equation'")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
import re

--fcc-editable-region--
class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')
            else:
                terms.append(f"{coefficient:+}x**{n}")
        equation_string = ' '.join(terms) + ' = 0'
        return re.sub(r"(?<!\d)1(?=x)", "", equation_string.strip("+"))        

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return [x]

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}

class QuadraticEquation(Equation):
    degree = 2
--fcc-editable-region--
    def __init__(self, *args):
        super().__init__(*args)
        a, b, c = self.coefficients.values()
        self.delta = b**2 - 4 * a * c

    def solve(self):
        if self.delta < 0:
            return []
        a, b, _ = self.coefficients.values()
        x1 = (-b + (self.delta) ** 0.5) / (2 * a)
        x2 = (-b - (self.delta) ** 0.5) / (2 * a)
        if self.delta == 0:
            return [x1]

        return [x1, x2]

    def analyze(self):
        a, b, c = self.coefficients.values()
        x = -b / (2 * a)
        y = a * x**2 + b * x + c
        if a > 0:
            concavity = 'upwards'
            min_max = 'min'
        else:
            concavity = 'downwards'
            min_max = 'max'
        return {'x': x, 'y': y, 'min_max': min_max, 'concavity': concavity}

def solver(equation):
    if not isinstance(equation, Equation):
        raise TypeError("Argument must be an Equation object")


lin_eq = LinearEquation(2, 3)
quadr_eq = QuadraticEquation(1, 2, 1)

```
