---
id: 6675aaf418b41157f6ccd692
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

現在是回到 `__init__` 方法的時候了。根據方程式型別，你在具現化實體時需要傳遞可變數量的引數。

為你的方法添加第二個參數 `args`，並使用 `*` 運算子使其接受可變數量的引數。

# --hints--

你的 `__init__` 方法應該接受兩個參數，`self` 和 `*args`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self, *args")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    degree = 1

    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
