---
id: 66759e32b88fb5459b1e0234
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

每當定義 `__init_subclass__` 的類別被繼承時，會呼叫 `__init_subclass__` 方法，並且它可用來自訂子類別。該方法接受一個慣例命名為 `cls`（代表「類別」）的參數，該參數表述新的子類別。

在你的 `Equation` 類別中定義一個 `__init_subclass__` 方法，並給它一個 `cls` 參數。

# --hints--

你應該在你的 `Equation` 類別中定義一個帶有 `cls` 參數的 `__init_subclass__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").has_args("cls")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    
--fcc-editable-region--
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
