---
id: 66793c5b4bdacc17c40ff8e7
title: 步驟 52
challengeType: 20
dashedName: step-52
---

# --description--

f-string 的另一個特性啟用你使用 `!` 後接轉換型別 `s`，將取代欄位（大括號）的內容轉換成字串。<dfn>例如</dfn>，`f'{obj!s}'` 將 `obj` 轉換成字串，且它等同於 `f'{str(obj)}'`。

從現在開始，你將持續透過將字串串接到 `output_string` 來建置輸出。

產生一個字串，該字串包含你的方程式的字串表述，並置中於寬度為 `24` 字元。讓字串以兩個換行字元開始和結束，並將你的新字串添加到 `output_string` 的當前值中。

`solver` 函式（程式）的輸出應該看起來像這樣：

```md

----Linear Equation-----

       2x +3 = 0        

```

# --hints--

`solver` 函式（程式）應該傳回不同的字串。

```js
({ test: () => runPython(`
eq = LinearEquation(4, 3)
expected = """
----Linear Equation-----

       4x +3 = 0        

"""
assert solver(eq) == expected
`) })
```

你應該使用轉換旗標 `!s` 將 `equation` 轉換成字串。

```js
({ test: () => (runPython(`
function = str(_Node(_code).find_function("solver"))
assert "equation!s" in function
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
import re


class Equation(ABC):
    degree: int
    type: str
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
        if not hasattr(cls, "type"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'type'"
            )
            
    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')
            else:
                terms.append(f"{coefficient:+}x**{n}")
        equation_string = ' '.join(terms) + ' = 0'
        return re.sub(r"(?<!\d)1(?=x)", "", equation_string.strip("+"))        

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    degree = 1
    type = 'Linear Equation'
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return [x]

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}


class QuadraticEquation(Equation):
    degree = 2
    type = 'Quadratic Equation'

    def __init__(self, *args):
        super().__init__(*args)
        a, b, c = self.coefficients.values()
        self.delta = b**2 - 4 * a * c

    def solve(self):
        if self.delta < 0:
            return []
        a, b, _ = self.coefficients.values()
        x1 = (-b + (self.delta) ** 0.5) / (2 * a)
        x2 = (-b - (self.delta) ** 0.5) / (2 * a)
        if self.delta == 0:
            return [x1]

        return [x1, x2]

    def analyze(self):
        a, b, c = self.coefficients.values()
        x = -b / (2 * a)
        y = a * x**2 + b * x + c
        if a > 0:
            concavity = 'upwards'
            min_max = 'min'
        else:
            concavity = 'downwards'
            min_max = 'max'
        return {'x': x, 'y': y, 'min_max': min_max, 'concavity': concavity}


def solver(equation):
    if not isinstance(equation, Equation):
        raise TypeError("Argument must be an Equation object")
--fcc-editable-region--
    output_string = f'\n{equation.type:-^24}'
    
--fcc-editable-region--
    return output_string

lin_eq = LinearEquation(2, 3)
quadr_eq = QuadraticEquation(1, 2, 1)
print(solver(lin_eq))

```
