---
id: 65f055b9190fc41ca35549b8
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Python пропонує різні методи, які включають у свої назви подвійне підкреслення на початку та в кінці. Можливо, ви вже знайомі з деякими, наприклад `__init__` та `__str__`. Ці методи, які відповідають шаблону імен `__<name>__`, називаються спеціальними методами, магічними методами або *дандер* (що англійською означає подвійне підкреслення(dunder (double underscore))).

Визначення спеціальних методів у класі впливає на поведінку цього класу. Вони викликаються під капотом у певних ситуаціях (наприклад, `__init__` під час створення екземпляра, `__str__`, коли об’єкт друкується або передається до `str()`). У цьому проєкті ви дізнаєтесь про деякі з найбільш часто використовуваних спеціальних методів.

Наразі призначте `x` до атрибута `x` об’єкта `Vector`.

# --hints--

Призначте `x` до `self.x`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").has_stmt("self.x = x")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        pass
--fcc-editable-region--
```
