---
id: 65f93b67169a9c703264458a
title: Крок 30
challengeType: 20
dashedName: step-30
---

# --description--

Оскільки метод має повернути рядок для створення екземпляра об’єкта для `R2Vector`, а також для `R3Vector` у разі успадкування, ви не можете будувати рядок із зазначенням імені класу.

Ви можете отримати доступ до імені класу за допомогою `__class__.__name__`. Додайте інструкцію `return` до методу `__repr__` та поверніть рядок, необхідний для створення екземпляра об’єкта.

# --hints--

Використайте `__class__.__name__` для будування f-рядка, необхідного для створення екземпляра об’єкта вектора та поверніть його з методу `__repr__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_return("f'{self.__class__.__name__}({args})'")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
