---
id: 65f9bc10a9fb1612e066e9e8
title: Крок 42
challengeType: 20
dashedName: step-42
---

# --description--

Коли додаються два вектори кожен компонент одного вектора додається до того самого компонента іншого вектора. Наприклад, додавання `(1, 2)` та `(2, 4)` створює третій вектор `(3, 6)`, де `3` — це сума `1` та `2`, а `6` — це сума `2` та `4`.

Словник `kwargs` міститиме пари ключ-значення, необхідні для створення нового екземпляру вектора того самого класу як і у двох векторів, складених разом.

Перетворіть порожній словник на розуміння словника, який ітерує над `vars(self)` та для кожного ключа (`i`) створює пару ключ-значення, де ключ є тим самим ключем поточної ітерації, а його значення є сумою значень атрибутів `i` двох операндів.

# --hints--

Перетворіть порожній словник, призначений до `kwargs`, на розуміння словника, яке ітерує над `vars(self)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_iters()[0].is_equivalent("vars(self)")
    `))
})
```

Розуміння словника, призначене до `kwargs`, має використовувати змінну `i`, щоб ітерувати над `vars(self)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

Розуміння словника, призначене до `kwargs`, має використовувати `i` як ключ.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_key().is_equivalent("i")
    `))
})
```

Розуміння словника, призначене `kwargs`, має призначати суму атрибута `i` з `self` та атрибута `i` з `other` до ключа `i` для кожного `i` у `vars(self)`. Для цього використайте функцію `getattr()`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) + getattr(other, i)") or node.is_equivalent("getattr(other, i) + getattr(self, i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {}
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
