---
id: 65f9b17a638f0d0dcce8c354
title: Крок 38
challengeType: 20
dashedName: step-38
---

# --description--

Щоб створити векторний простір, вам потрібно визначити, як вектори повинні поводитися в кількох випадках. Вектори можна додавати, утворюючи новий вектор.

Спеціальний метод `__add__` може бути реалізований для перевизначення того, що відбувається за замовчуванням, коли два об’єкти додаються разом за допомогою оператора `+`.

Прямо зараз спроба додати два екземпляри `R2Vector` або `R3Vector` призведе до підняття виключення. Створіть порожній метод `__add__` в межах класу `R2Vector` і надайте йому два параметри: `self` та `other`.

# --hints--

Визначте метод `__add__` в межах класу `R2Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__add__")`)) })
```

Метод `__add__` повинен мати два параметри: `self` та `other`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__add__").has_args("self, other")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--

--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
