---
id: 65faaed8f7a9772f023ea816
title: Крок 52
challengeType: 20
dashedName: step-52
---

# --description--

Вектори можна помножити на *скаляр*, тобто число, яке множить кожен окремий компонент. Результатом скалярного множення є вектор з такою ж орієнтацією, що й оригінальний вектор, але іншої величини.

Застосуйте скалярне множення, перевіривши, чи `other` є `int` або `float`. Якщо так, поверніть новий екземпляр поточного класу, який має кожен компонент початкового вектора помноженим на `other`. Це буде вектор, отриманий у результаті скалярного множення.

Переконайтеся, що ці методи можна застосувати для обчислення скалярного множення векторів із будь-якою кількістю вимірів.

# --hints--

Метод має повертати новий екземпляр поточного класу лише тоді, коли тип `other` є `int` або `float`.

```js
({ test: () => runPython(`
v1 = R2Vector(x=0, y=0)
v2 = R3Vector(x=0, y=0, z=0)
types = ["", True, [], {}]
assert all((v1 * i) is None for i in types)
types = [1, 1.0]
assert all(type(v1 * i) is type(v1) for i in types)
assert all(type(v2 * i) is type(v2) for i in types)
`) })
```

Вектор, отриманий в результаті скалярного множення, повинен мати кожен компонент з початкового вектора помноженим на скаляр.

```js
({ test: () => runPython(`
v1 = R2Vector(x=1, y=1.5)
v2 = R3Vector(x=2.2, y=3, z=-1)
assert vars(v1 * 3) == vars(R2Vector(x=3, y=4.5))
assert vars(v2 * (-2.0)) == vars(R3Vector(x=-4.4, y=-6.0, z=2.0))
`) })
```


# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)

    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) - getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    def __mul__(self, other):
        pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R2Vector(x=0.5, y=1.25)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
v3 = v1 + v2
print(f'v1 + v2 = {v3}')
v4 = v1 - v2
print(f'v1 - v2 = {v4}')

```
