---
id: 666827a6fd0dbaafe8330ea6
title: Крок 35
challengeType: 20
dashedName: step-35
---

# --description--

`__getattr__` — ще один спеціальний метод, який відіграє роль у доступі до атрибутів.

Реалізація `__getattribute__` за замовчуванням це підняти `AttributeError`, коли запитуваний атрибут не є атрибутом екземпляра або його немає в дереві класів.

У цьому випадку `__getattr__` викликається, якщо визначено класом. Ви можете розглядати це як свого роду запасний варіант, коли звичайна процедура доступу до атрибутів падає.

Перетворіть метод `__getattribute__` на `__getattr__` та змініть рядок, який повертає метод, на `'calling __getattr__'`.

# --hints--

Визначте метод під назвою `__getattr__` з двома параметрами, `self` та `attr`, в межах класу `R2Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattr__").has_args("self, attr")`)) })
```

Метод `__getattr__` має повернути рядок `'calling __getattr__'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattr__").has_return("'calling __getattr__'")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __getattribute__(self, attr):                
        return 'calling __getattribute__'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')

print(v1.x)
print(getattr(v1, 'x'))

```
