---
id: 66682150151af29efec9727d
title: Крок 33
challengeType: 20
dashedName: step-33
---

# --description--

Метод `__getattribute__` викликається під капотом кожного разу, коли ви намагаєтеся отримати доступ до атрибута екземпляра. Якщо атрибут не знайдено на рівні екземпляра, метод шукатиме його на рівні класу та, зрештою, у батьківських класах.

В межах класу визначте метод `__getattribute__` з двома параметрами, `self` та `attr`, і змусьте його повертати рядок `'calling __getattribute__'`. Ви миттєво заміните реалізацію за замовчуванням, тільки щоб побачити, як працює доступ до атрибутів.

# --hints--

Визначте метод під назвою `__getattribute__` з двома параметрами, `self` та `attr`, в межах класу `R2Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_args("self, attr")`)) })
```

Метод `__getattribute__` має повернути рядок `'calling __getattribute__'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_return("'calling __getattribute__'")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--

--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')

```
