---
id: 56533eb9ac21ba0edf2244b3
title: Construir um Conversor de Celsius para Fahrenheit
challengeType: 26
dashedName: build-a-celsius-to-fahrenheit-converter
---

# --description--

Neste laboratório, você escreverá uma função que converte a temperatura de Celsius para Fahrenheit.A fórmula para converter de Celsius para Fahrenheit é:

```js
fahrenheit = celsius * (9/5) + 32
```

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma função chamada `convertCtoF`.
2. `convertCtoF` deve receber um único argumento numérico, que é a temperatura em Celsius.
3. `convertCtoF` deve retornar um número.

# --hints--

Você deve criar uma função chamada `convertCtoF`.

```js
assert.isFunction(convertCtoF);
```

`convertCtoF` deve receber um único parâmetro.

```js
assert.lengthOf(convertCtoF, 1);
```

`convertCtoF(0)` deve retornar um número.

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` deve retornar um valor de `-22`.

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` deve retornar um valor de `14`.

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` deve retornar um valor de `32`.

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` deve retornar um valor de `68`.

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` deve retornar um valor de `86`.

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
