---
id: 65ef198fde24dfb7ff675b42
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

`raise` 文を使うと、特定の例外を強制的に発生させることができます。これは `raise` キーワードに続けて例外の種類を指定し、カスタムのエラーメッセージを提供できます。

```py
raise ValueError("Invalid value")
```

上記のコードが実行されると、`ValueError` が発生し、メッセージ `"Invalid value"` がユーザーに表示されます。

もし `square_target` が `0` より小さい場合、実数の平方根を計算できません。したがって、メッセージ `ValueError` を付けて `'Square root of negative number is not defined in real numbers'` を発生させてください。`pass` キーワードを削除するのを忘れないでください。

# --hints--

`pass` キーワードを削除してください。

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].has_pass()`))
    }
})
```

`ValueError` の中で、メッセージ `'Square root of negative number is not defined in real numbers'` を付けて `if` を発生させてください。

```js
({
    test: () => {
        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].is_equivalent("raise ValueError('Square root of negative number is not defined in real numbers')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        pass

--fcc-editable-region--
```
