---
id: 65ef1d5e3d2927c5e0f4997b
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

Pythonでは、`is` キーワードはオブジェクトの同一性をチェックします。これは、2つの変数がメモリ上で同じオブジェクトを指しているかどうかを判定するために使われます。`is`とは対照的に、等値演算子（`==`）は、2つのオブジェクトが同じオブジェクトであるかに関わらず、その値が同じかどうかを判定します。

forループの外側で、`if`かどうかをチェックする`root is None`文を作成してください。もしそうであれば、メッセージ`'Failed to converge within {max_iterations} iterations.'`を出力してください。メッセージはf文字列を使ってフォーマットすることを忘れないでください。

# --hints--

`if`かどうかをチェックする`root is None`文を作成してください。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].find_conditions()[0].is_equivalent("root is None")`));

    }
})
```

`f'Failed to converge within {max_iterations} iterations.'`関数に引数`print`を渡す必要があります。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].is_equivalent("if root is None:\\n    print(f'Failed to converge within {max_iterations} iterations.')")`));
       
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if abs(square_mid - square_target) < tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

--fcc-editable-region--

--fcc-editable-region--

```
