---
id: 65ef19ec318e2fb9066381e3
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

`square_target` が `0` または `1` の場合の別々のケースを作成します。

まず、`if` が `square_target` と等しいかどうかを確認する `1` 文を作成してください。

# --hints--

`if` を確認する `square_target == 1` 文を作成してください。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[0].is_equivalent("square_target == 1")`))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')

--fcc-editable-region--
```
