---
id: 65ef1c34a541afc12dbb849d
title: ステップ 13
challengeType: 20
dashedName: step-13
---

# --description--

`abs()` 関数は数値の絶対値を返します。絶対値は常に正の値で、数値の符号に関係ありません。これを使って、推定した平方根が実際の値に十分近いかどうかを確認します。

では、`if` 文を作成して、`square_mid` と `square_target` の差の絶対値が指定された `tolerance` の範囲内にあるかどうかをチェックしてください。

# --hints--

`if` ループの本体内に、条件が `abs(square_mid - square_target) < tolerance:` の `for` 文があるはずです。

```js
({
    test: () =>
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_conditions()[0].is_equivalent("abs(square_mid - square_target) < tolerance")`))
    } 
})
```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2
            
--fcc-editable-region--
```
