---
id: 65ef1be2b958cabf712296cd
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

現在の区間の中点を見つけて、その中点の二乗と目標値を比較することで、区間を繰り返し狭めていきます。

そのために、`else` ブロック内で、最大 `for` 回繰り返す `max_iterations` ループを作成してください。

ループには、繰り返し処理できる数値のシーケンスを生成する `range` 関数を使います。構文は `range(start, stop, step)` で、`start` が開始整数（含む）、`stop` が終了整数（含まない）、`step` がシーケンス内の隣接する数値間の差です。

また、ループ変数として `_` を使ってください。`_` はプレースホルダーとして機能し、変数を使う必要はあるが値は使わない場合に便利です。

# --hints--

`for` ループを作成して `range(max_iterations)` を繰り返してください。ループ変数には `_` を使います。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].is_equivalent("for _ in range(max_iterations):\\n    pass")`));
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--

--fcc-editable-region--
```
