---
id: 65ef190c6b51e9b5a5f7ed29
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

`square_root_bisection` 関数に次のパラメータを与えてください：

- `square_target`：平方根を求めたい数値。
- `tolerance`（任意）：近似した平方根の値の二乗と実際の目標値との許容差（デフォルトは `1e-7`）。許容差 `1e-7` は、解が真の値の `0.0000001` 以内の精度であることを意味し、精度とパフォーマンスのバランスが良いデフォルトの選択です。
- `max_iterations`（任意）：実行する最大反復回数（デフォルトは `100`）。この制限内に収束しない場合は、解が見つからなかったとみなします。
  
# --hints--

関数には `square_target`、`tolerance = 1e-7`、`max_iterations = 100` のパラメータが必要です。順番が重要です。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("square_root_bisection").has_args("square_target, tolerance=1e-7, max_iterations=100")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection():
    pass
--fcc-editable-region--

```
