---
id: 65ef19425d1b27b6c930bae6
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

平方根を求めたい数が負の数の場合、負の数の平方根は実数では定義されていないため、コードはエラーを発生させる必要があります。

`pass` の文を削除し、`if` が `square_target` より小さいかどうかを確認する `0` の文を作成してください。

# --hints--

`pass` キーワードを削除してください。

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").has_pass()`))
    }
})
```

`if` を確認する `square_target < 0` の文を作成してください。

```js
({
    test: () => {
       assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_if("square_target < 0").is_empty()`))

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance = 1e-7, max_iterations = 100):
    pass
--fcc-editable-region--
```
