---
id: 65ef1a50049cf9bada13266f
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

`elif` 文を作成して、`square_target` が `0` と等しいかどうかを確認してください。等しい場合は、`0` の値を `root` 変数に代入してください。また、メッセージ `'The square root of {square_target} is 0'` を出力してください。メッセージは f-string を使ってフォーマットすることを忘れないでください。

# --hints--

条件 `elif` を確認するために `square_target == 0` 文が必要です。


```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[1].is_equivalent("square_target == 0")`));
    }
})

```

値 `0` を `root` 変数に代入し、引数 `f'The square root of {square_target} is 0'` を `print` の呼び出しに渡す必要があります。


```js

({
    test: () => 
    {        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[1].is_equivalent("root = 0\\nprint(f'The square root of {square_target} is 0')")`));
    }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    
--fcc-editable-region--
```
