---
id: 65ef181dad3d96b14a21a03e
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このプロジェクトでは、二分法を使って与えられた数の近似平方根を求めます。

二分法は実数値関数の根を見つける手法です。平方根が存在する区間を徐々に狭めていき、指定された許容誤差内の値に収束させます。

まず、`square_root_bisection` という名前の関数を作成してください。パラメータはまだ空のままにしておきます。

# --hints--

パラメータなしで `square_root_bisection` 関数を定義してください。

```js
({
    test: () => {
        assert(runPython(`
        import inspect        
        inspect.isfunction(square_root_bisection)
        `))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
