---
id: 69df434b56d2e0faf1a7424c
title: Schritt 25
challengeType: 0
dashedName: step-25
---

# --description--

Da Ihre fCC-Autoren-Seite jetzt vollständig funktionsfähig ist, wollen wir sie umgestalten, um die Lesbarkeit zu verbessern, indem wir `async` und `await` anstelle der Promise-Verkettungsmethode `then` verwenden.

Denken Sie daran, dass Sie, um den `await`-Operator zu verwenden, um auf eine Funktion zu warten, die ein `Promise` zurückgibt, den Funktionsaufruf in eine `async`-Funktion einwickeln müssen (wenn sie sich nicht im Hauptteil eines Moduls befindet).

Da Ihr `fetch`-Aufruf weder in einer asynchronen Funktion noch im Hauptteil eines Moduls definiert ist, können Sie `await` erst verwenden, wenn Sie das behoben haben.

Verpacken Sie die `fetch`-Anweisung und die gesamte Folge der verketteten then- und catch-Methoden als Ganzes in eine neue asynchrone Pfeilfunktion namens `initialFetch`, die keine Argumente annimmt.

Ein Beispiel für eine `async`-Funktionsdeklaration ist:

```js
const functionName = async () => {}
```

# --before-all--

```js
window.fetch = () => Promise.resolve({json: () => Promise.resolve([{ author: 'Whoever', image: 'data:image/gif;base64,R0lGODlhAQABAIAAAAAAAP///ywAAAAAAQABAAACAUwAOw==', url: "http://not-a-real-url.nowhere/", bio: 'words go here' }])});
```    

# --hints--

`initialFetch` sollte eine Funktion sein.

```js
assert.isFunction(initialFetch)
```

Sie sollten `const` verwenden, um eine `initialFetch`-Funktion zu erstellen.

```js
assert.match(code, /const\s+initialFetch\s*=\s*/)
```

`initialFetch` sollte eine asynchrone Funktion sein.

```js
assert.isFunction(initialFetch)
assert.match(code, /const\s+initialFetch\s*=\s*async\s*/)
```

Ihre `initialFetch`-Funktion sollte keine Parameter annehmen.

```js
assert.match(code, /const\s+initialFetch\s*=\s*.*\s*\(\s*\)\s*/)
```

Ihre `initialFetch`-Funktion sollte die Pfeilsyntax verwenden.

```js
assert.match(code, /const\s+initialFetch\s*=\s*.*\s*\(.*\)\s*=>\s*/)
```

Ihre `initialFetch`-Funktion sollte nicht leer sein.

```js
assert.isFunction(initialFetch)
assert.notMatch(code, /const\s+initialFetch\s*=\s*.*\s*\(\s*\)\s*=>\s*\{\s*\}/)
```

Ihre `initialFetch`-Funktion sollte die `fetch`-Anweisung und die verketteten Methoden enthalten.

```js
const initialFetchBlockRegex =
  /const\s+initialFetch\s*=\s*async\s*\(\s*\)\s*=>\s*\{([\s\S]*?)\}\s*$/m;

const match = code.match(initialFetchBlockRegex);

assert.isDefined(match);

const initialFetchBlockContents = match[1].trim();
const fetchStatementRegex = /^fetch\(\s*(['"`])https:\/\/cdn\.freecodecamp\.org\/curriculum\/news-author-page\/authors\.json\1\s*\)\s*\.then\(\s*(?:\(\s*res\s*\)|res)\s*=>\s*res\.json\(\s*\)\s*\)\s*\.then\(\s*(?:\(\s*data\s*\)|data)\s*=>\s*\{\s*authorDataArr\s*=\s*data;?\s*displayAuthors\(\s*authorDataArr\.slice\(\s*startingIndex\s*,\s*endingIndex\s*\)\s*\)\s*;?\s*\}\s*\)\s*\.catch\(\s*(?:\(\s*err\s*\)|err)\s*=>\s*\{?\s*authorContainer\.innerHTML\s*=\s*(['"`])<p\s+class\s*=\s*(['"`])error-msg\3>\s*There\s+was\s+an\s+error\s+loading\s+the\s+authors\s*<\/p>\s*\2\s*;?\s*\}?\s*\)\s*;?/;

assert.isTrue(fetchStatementRegex.test(initialFetchBlockContents));

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

let startingIndex = 0;
let endingIndex = 8;
let authorDataArr = [];

--fcc-editable-region--

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    authorDataArr = data;
    displayAuthors(authorDataArr.slice(startingIndex, endingIndex));
  })
  .catch((err) => {
    authorContainer.innerHTML = '<p class="error-msg">There was an error loading the authors</p>';
  });

--fcc-editable-region--

const fetchMoreAuthors = () => {
  startingIndex += 8;
  endingIndex += 8;

  displayAuthors(authorDataArr.slice(startingIndex, endingIndex));
  if (authorDataArr.length <= endingIndex) {
    loadMoreBtn.disabled = true;
    loadMoreBtn.style.cursor = 'not-allowed';
    loadMoreBtn.textContent = 'No more data to load';
  }
};

const displayAuthors = (authors) => {
  authors.forEach(({ author, image, url, bio }, index) => {
    authorContainer.innerHTML += `
      <div id="${index}" class="user-card">
        <h2 class="author-name">${author}</h2>
        <img class="user-img" src="${image}" alt="${author} avatar" />
        <div class="purple-divider"></div>
        <p class="bio">${bio.length > 50 ? bio.slice(0, 50) + '...' : bio}</p>
        <a class="author-link" href="${url}" target="_blank">${author}'s author page</a>
      </div>
    `;
  });
};

loadMoreBtn.addEventListener('click', fetchMoreAuthors);
```
