---
id: 66c645b345a39bf04864dc50
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

업데이트된 쇼핑 목록을 확인하려면 현재 쇼핑 목록과 간단한 메시지를 콘솔에 출력해야 합니다.

이 메시지는 워크숍 내내 반복해서 사용할 것이므로 재사용 가능한 함수를 만드는 것이 가장 좋습니다.

배열을 매개변수로 받고, 제공된 배열의 내용에 이어서 문자열 `getShoppingListMsg`을 반환하는 `"Current Shopping List: "`라는 함수를 만드세요.

템플릿 리터럴이나 `+` 연산자를 사용해 문자열과 함수에 전달된 배열을 결합할 수 있습니다.

# --hints--

코드에 `getShoppingListMsg`라는 함수가 있어야 합니다.

```js
assert.isFunction(getShoppingListMsg);
```

`getShoppingListMsg` 함수는 문자열을 반환해야 합니다.

```js
assert.isString(getShoppingListMsg(shoppingList));
```

`getShoppingListMsg` 함수는 배열을 받아서 그 배열 내용에 이어서 메시지 `"Current Shopping List: "`를 반환해야 합니다. 띄어쓰기와 철자 오류를 다시 확인하세요.

```js
assert.oneOf(getShoppingListMsg(["Anchovies", "Olives"]), ["Current Shopping List: Anchovies,Olives", "Current Shopping List: Anchovies, Olives"]);
assert.strictEqual(getShoppingListMsg(["Apples"]), "Current Shopping List: Apples");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

--fcc-editable-region--

--fcc-editable-region--
```
