---
id: 66c63bfa4685e4e3843fa0dc
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 배열에 대해 계속 배우면서 식료품 쇼핑 목록을 만들어 보세요.

먼저, 문자열 `console.log`을 콘솔에 기록하는 `"Grocery shopping list"`를 추가하세요.

# --hints--

코드에 `console.log`가 있어야 합니다.

```js
assert.match(code, /console\.log(.*)/);
```

`console.log`는 메시지 `"Grocery shopping list"`를 출력해야 합니다. 철자와 대소문자 오류가 없는지 다시 확인하세요.

```js
assert.match(code, /console\.log\((["'])(Grocery\s+shopping\s+list)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
