---
id: 66c6491fe4c8e0f16845425f
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

이제 목록에 또 다른 과일을 추가할 차례입니다.

앞서 사용한 것과 같은 배열 메서드를 사용해 `"Grapes"` 배열의 끝에 문자열 `shoppingList`을 추가하세요.

그다음 `console.log`를 추가하고 `getShoppingListMsg` 안에서 `shoppingList` 배열을 인수로 하여 `console.log` 함수를 호출해 업데이트된 목록이 콘솔에 출력되도록 하세요.

# --hints--

`shoppingList` 배열에는 이제 두 개의 항목이 있어야 합니다.

```js
assert.lengthOf(shoppingList, 2);
```

`shoppingList` 배열에는 문자열 `"Grapes"`이 포함되어야 합니다.

```js
assert.equal(shoppingList[1], "Grapes");
```

`getShoppingListMsg(shoppingList)` 안에서 `console.log` 함수를 호출해야 합니다.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
