---
id: 66c64095f61166ede6e39a84
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

이 쇼핑 목록에서는 구매해야 할 항목을 나타내기 위해 배열을 사용하세요.

이전 수업에서 다음과 같이 배열을 만드는 방법을 배웠습니다:

```js
const fruits = ['apple', 'banana', 'orange'];
```

이번 단계에서는 `shoppingList`라는 변수를 만들고 빈 배열을 할당하세요.

# --hints--

`shoppingList`라는 변수가 있어야 합니다.

```js
assert.isNotNull(shoppingList);
```

`shoppingList` 변수에 배열을 할당해야 합니다.

```js
assert.isArray(shoppingList);
```

배열은 비어 있어야 합니다.

```js
assert.isEmpty(shoppingList);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

--fcc-editable-region--

--fcc-editable-region--
```
