---
id: 66c64210413532ee9d3bd342
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

이 워크숍의 다음 부분에서는 식료품 목록에 음식 항목을 추가하는 연습을 하세요.

먼저 `console.log`를 사용해 `"It will be nice to have some fruit to eat."` 메시지를 기록하세요.

# --hints--

콘솔에 `"It will be nice to have some fruit to eat."` 메시지를 기록해야 합니다.

```js
assert.match(code, /console\.log\s*\(\s*(['"])(It\s+will\s+be\s+nice\s+to\s+have\s+some\s+fruit\s+to\s+eat\.)\1\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

--fcc-editable-region--

--fcc-editable-region--
```
