---
id: 66c644c08ebcd2ef90c750bd
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

이전 수업에서 `push` 메서드를 사용해 배열 끝에 항목을 추가하는 방법을 배웠습니다:

```js
fruits.push('pear');
```

`push` 메서드를 사용해 `"Apples"` 배열에 문자열 `shoppingList`를 추가하세요.

# --hints--

`shoppingList` 배열에는 항목이 하나 있어야 합니다.

```js
assert.lengthOf(shoppingList, 1);
```

`shoppingList` 배열에는 문자열 `"Apples"`가 있어야 합니다. 철자와 대소문자를 다시 확인하세요.

```js
assert.includeMembers(shoppingList, ['Apples']);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

--fcc-editable-region--

--fcc-editable-region--
```
