---
id: 66c64ad16796c7f2419b45c5
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

이제 장보기 목록 상단에 항목을 추가할 시간입니다.

먼저 `console.log`를 추가하여 메시지 `"It looks like we need to get some cooking oil."`을 콘솔에 기록하세요.

# --hints--

메시지 `"It looks like we need to get some cooking oil."`를 콘솔에 기록해야 합니다.

```js
assert.match(code, /console\.log\((['"])(It\s+looks\s+like\s+we\s+need\s+to\s+get\s+some\s+cooking\s+oil\.)\1\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
