---
id: 66c730ee6ae076281721d0b9
title: 14단계
challengeType: 1
dashedName: step-14
---

# --description--

이전 수업에서 `pop` 메서드를 사용해 배열 끝에서 항목을 제거하는 방법을 배웠습니다.

다음은 `pop` 메서드를 사용하는 방법에 대한 복습입니다:

```js
let array = [1, 2, 3, 4, 5];
array.pop();

// [1, 2, 3, 4]
console.log(array); 
```

`pop` 메서드를 사용해 `shoppingList` 배열의 마지막 항목을 제거하세요.

# --hints--

`shoppingList` 배열에는 이제 `5` 항목이 있어야 합니다.

```js
assert.lengthOf(shoppingList, 5);
```

`shoppingList` 배열에는 더 이상 문자열 `"Potato Chips"` 이 없어야 합니다.

```js
assert.notInclude(shoppingList, "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

--fcc-editable-region--

--fcc-editable-region--
```
