---
id: 66c73a7798f6f62b2ae58f22
title: 16단계
challengeType: 1
dashedName: step-16
---

# --description--

이제 식료품 목록의 맨 앞에 항목을 더 추가할 시간입니다.

먼저 `console.log` 문장을 추가하여 `"It might be nice to get a dessert."` 메시지를 기록하세요.

그 아래 `console` 문장에는 올바른 배열 메서드를 사용해 `"Chocolate Cake"` 문자열을 `shoppingList` 배열의 맨 앞에 추가하세요.

마지막으로 `console.log`를 추가하고 `getShoppingListMsg` 안에서 `shoppingList` 배열을 인수로 하여 `console.log` 함수를 호출해 업데이트된 목록이 콘솔에 기록되는 것을 확인하세요.

# --hints--

`console.log`가 `"It might be nice to get a dessert."` 메시지를 기록해야 합니다.

```js
assert.match(code, /console\.log\((['"])(It\s+might\s+be\s+nice\s+to\s+get\s+a\s+dessert\.)\1\)/);
```

`shoppingList` 배열에는 총 `6`개의 항목이 있어야 합니다.

```js
assert.lengthOf(shoppingList, 6);
```

`unshift` 배열의 맨 앞에 `"Chocolate Cake"` 문자열을 추가하려면 `shoppingList` 메서드를 사용하세요.

```js
assert.strictEqual(shoppingList[0], "Chocolate Cake");
```

`getShoppingListMsg(shoppingList)` 안에서 `console.log`를 호출해야 합니다.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 6);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
