---
id: 66c73a0c5b264f2a75164d94
title: 15단계
challengeType: 1
dashedName: step-15
---

# --description--

이제 업데이트된 `shoppingList` 배열을 콘솔에 기록할 차례입니다.

`console.log` 안에 `getShoppingListMsg`을 추가하고 `shoppingList` 배열을 인수로 하여 `console.log` 함수를 호출해 업데이트된 목록이 콘솔에 기록되는 것을 확인하세요.

# --hints--

`getShoppingListMsg(shoppingList)` 안에서 `console.log`를 호출해야 합니다.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 5);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();

--fcc-editable-region--

--fcc-editable-region--
```
