---
id: 66c72a55418cc9247b710827
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

다음으로, `console.log` 안에서 `getShoppingListMsg` 배열을 인수로 하여 `shoppingList` 함수를 호출하고 `console.log`을 추가해 업데이트된 목록이 콘솔에 출력되도록 하세요.

# --hints--

`getShoppingListMsg(shoppingList)` 안에서 `console.log`를 호출해야 합니다.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 3);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");

--fcc-editable-region--

--fcc-editable-region--
```
