---
id: 66c74079c30b1c2d166cb9a4
title: 18단계
challengeType: 1
dashedName: step-18
---

# --description--

이전 수업에서 `shift` 메서드를 사용해 배열의 첫 항목을 제거하는 방법을 배웠습니다.

다음은 `shift` 메서드를 사용하는 방법에 대한 복습입니다:

```js
const array = [1, 2, 3, 4, 5];
array.shift();

// Result: [2, 3, 4, 5]
console.log(array); 
```

`shift` 메서드를 사용해 `shoppingList` 배열의 첫 번째 항목을 제거하세요.

# --hints--

`shoppingList` 배열에는 총 `5` 개의 항목이 있어야 합니다.

```js
assert.lengthOf(shoppingList, 5);
```

`shift` 메서드를 사용해 `shoppingList` 배열의 첫 번째 항목을 제거해야 합니다.

```js
assert.strictEqual(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

console.log("It might be nice to get a dessert.");

shoppingList.unshift("Chocolate Cake");
console.log(getShoppingListMsg(shoppingList));

console.log("On second thought, maybe we should be more health conscious.");

--fcc-editable-region--

--fcc-editable-region--
```
