---
id: 66c742d045c9fc2e09fa64b1
title: 19단계
challengeType: 1
dashedName: step-19
---

# --description--

장보기 목록에서 마지막으로 변경할 부분은 목록의 첫 번째 항목을 업데이트하는 것입니다.

이전 수업에서 대괄호 표기법과 업데이트하려는 항목의 인덱스를 사용해 항목을 수정하는 방법을 배웠습니다.

배열에서 항목을 업데이트하는 방법을 다시 알려드립니다:

```js
const array = [1, 2, 3, 4, 5];

array[0] = 10;
// Result: [10, 2, 3, 4, 5]
console.log(array); 
```

`shoppingList` 배열의 첫 번째 항목을 `"Canola Oil"`로 업데이트하세요.

# --hints--

`shoppingList` 배열의 첫 번째 항목을 `"Canola Oil"`로 업데이트해야 합니다.

```js
assert.equal(shoppingList[0], "Canola Oil");
assert.notEqual(shoppingList[1], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

console.log("It might be nice to get a dessert.");

shoppingList.unshift("Chocolate Cake");
console.log(getShoppingListMsg(shoppingList));

console.log("On second thought, maybe we should be more health conscious.");

shoppingList.shift();

--fcc-editable-region--

--fcc-editable-region--
```
