---
id: 66c730183f4020275cbf0611
title: 13단계
challengeType: 1
dashedName: step-13
---

# --description--

이번 워크숍의 다음 부분에서는 배열 끝에서 항목을 제거하는 방법을 검토합니다.

먼저 `console.log`를 추가하여 `"This looks like too much junk food."` 메시지를 기록하세요.

# --hints--

`"This looks like too much junk food."` 메시지를 기록해야 합니다.

```js
assert.match(code, /console\.log\((['"])(This\s+looks\s+like\s+too\s+much\s+junk\s+food.)\1\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
