---
id: 66c72b0ffbc5522525768558
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

이전 단계에서 `push` 메서드를 사용해 배열 끝에 항목을 추가하는 방법을 검토했습니다.

하지만 `push` 메서드는 여러 인수를 받아서, 이렇게 배열 끝에 여러 항목을 추가할 수 있습니다:

```js
array.push(item1, item2, item3);
```

이번 단계에서는 `push` 메서드를 사용해 `"Popcorn", "Beef Jerky", "Potato Chips"` 문자열들을 `shoppingList` 배열에 추가하세요.

순서가 중요하니, 나열된 순서대로 항목을 추가해야 합니다.

# --hints--

`shoppingList` 배열에는 총 `6` 개의 문자열이 있어야 합니다.

```js
assert.lengthOf(shoppingList, 6);
```

`shoppingList` 에는 `"Popcorn"` 문자열이 있어야 합니다.

```js
assert.strictEqual(shoppingList[3], "Popcorn");
```

`shoppingList` 에는 `"Beef Jerky"` 문자열이 있어야 합니다.

```js
assert.strictEqual(shoppingList[4], "Beef Jerky");
```

`shoppingList` 에는 `"Potato Chips"` 문자열이 있어야 합니다.

```js
assert.strictEqual(shoppingList[5], "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
