---
id: 66c726c34ecf1e238aa9d7d5
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

이전 수업에서 `unshift()` 메서드를 사용해 배열의 시작 부분에 요소를 추가하는 방법을 배웠습니다.

다음은 `unshift()` 메서드를 사용하는 방법에 대한 복습입니다:

```js
array.unshift(item1, item2, ..., itemX);
```

`unshift()` 메서드를 사용해 `"Vegetable Oil"` 배열의 시작 부분에 문자열 `shoppingList` 을 추가하세요.

# --hints--

`shoppingList` 배열에 세 개의 항목이 있어야 합니다.

```js
assert.lengthOf(shoppingList, 3);
```

`unshift()` 메서드를 사용해 `"Vegetable Oil"` 배열의 시작 부분에 `shoppingList` 을 추가해야 합니다.

```js
assert.equal(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

--fcc-editable-region--

--fcc-editable-region--
```
