---
id: 66cbe2319d3845545a293a0b
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

이제 콘솔에 기록된 메시지를 확인할 차례입니다.

`console.log` 안에 `getShoppingListMsg` 배열을 인수로 하여 `shoppingList` 함수를 호출하고 `console.log`를 추가해 콘솔에 메시지가 기록되는 것을 확인하세요.

# --hints--

`getShoppingListMsg(shoppingList)` 안에서 `console.log` 함수를 호출해야 합니다.

```js
assert.match(code, /console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

--fcc-editable-region--

--fcc-editable-region--
```
