---
id: 60a3e3396c7b40068ad69994
title: Step 41
challengeType: 0
dashedName: step-41
---

# --description--

Die Eigenschaft `border-radius` akzeptiert bis zu vier Werte, um die Ecken oben links, oben rechts, unten rechts und unten links abzurunden.

Runden Sie die obere linke Ecke von `.three` um 30 Pixel ab, die obere rechte um 25 Pixel, die untere rechte um 60 Pixel und die untere linke um 12 Pixel.

# --hints--

Sie sollten die Eigenschaft `border-radius` auf `30px 25px 60px 12px` setzen.

```js
const hasBorderRadius = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['border-radius'] === '30px 25px 60px 12px');
assert(hasBorderRadius);
```

Ihr `.three`-Element sollte einen `border-radius`-Wert von `30px 25px 60px 12px` haben.

```js
const threeBorderRadius = new __helpers.CSSHelp(document).getStyle('.three')?.getPropertyValue('border-radius');
assert(threeBorderRadius === '30px 25px 60px 12px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
  filter: blur(2px);
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
  box-shadow: 0 0 3px 3px #efb762;
  border-radius: 9px;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
  margin: 0 auto 20px;
  box-shadow: 0 0 3px 3px #8f0401;
  border-radius: 8px 10px;
}

.one, .two {
  filter: blur(1px);
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
  filter: blur(2px);
  box-shadow: 0 0 5px 5px #b20403;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
