---
id: 60a3e3396c7b40068ad69977
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

Verwenden Sie Innenabstand, um den Abstand innerhalb eines Elements anzupassen.

Verwenden Sie in `.frame` die Kurzschreibweise-Eigenschaft `padding`, um den Abstand zwischen den `.frame`- und `.canvas`-Elementen um `50px` zu vergrößern. Die Kurzschreibweise erhöht den Abstand oben, unten, links und rechts am Rahmen und Canvas des Elements.

# --hints--

Sie sollten die Eigenschaft `padding` auf `50px` setzen.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

Ihr `.frame`-Element sollte einen `padding`-Wert von `50px` haben.

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
