---
id: 60a3e3396c7b40068ad6998b
title: Step 32
challengeType: 0
dashedName: step-32
---

# --description--

Es ist hilfreich, wenn Ihre Außenabstände in eine Richtung drücken.

In diesem Fall drückt der untere Außenabstand des `.one`-Elements `.two` um 20 Pixel nach unten.

Verwenden Sie im `.two`-Selektor die `margin`-Kurzschreibweise-Eigenschaft, um den oberen Außenabstand auf `0`, den horizontalen Außenabstand auf `auto` und den unteren Außenabstand auf `20px` zu setzen. Dadurch wird der obere Außenabstand entfernt, das Element horizontal zentriert und der untere Außenabstand auf 20 Pixel gesetzt.

# --hints--

Sie sollten die `margin`-Eigenschaft auf `0 auto 20px` setzen.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '0px auto 20px');
assert(hasMargin);
```

Ihr `.two`-Element sollte einen `margin`-Wert von `0 auto 20px` haben.

```js
const twoMargin = new __helpers.CSSHelp(document).getStyle('.two')?.getPropertyValue('margin');
assert(twoMargin === '0px auto 20px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
