---
id: 60a3e3396c7b40068ad69979
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

Fügen Sie ein neues `div`-Element innerhalb Ihres `.canvas`-Elements hinzu.

Geben Sie dem neuen `div` das `class`-Attribut mit dem Wert `one`. Dies wird Ihr erstes Rechteck sein.

# --hints--

Sie sollten ein neues `div`-Element erstellen.

```js
assert(document.querySelectorAll('div').length === 3);
```

Sie sollten das neue `div`-Element innerhalb Ihres `.canvas`-Elements verschachteln.

```js
assert(document.querySelector('.canvas').children[0].tagName === 'DIV');
```

Ihr neues `div` sollte ein `class`-Attribut mit dem Wert `one` haben.

```js
assert(document.querySelector('.canvas').children[0].className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
