---
id: 60a3e3396c7b40068ad6996d
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Der Inhalt ist von einem Raum umgeben, der Innenabstand genannt wird, ähnlich wie Luftpolsterfolie einen Gegenstand von der umgebenden Verpackung trennt.

Den Rahmen können Sie sich wie den Karton vorstellen, in dem Ihr Gegenstand verschickt wurde.

Ändern Sie das `src`-Attribut in `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

# --hints--

Das `img`-Element sollte ein `src` von `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png` haben

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```
