---
id: 60a3e3396c7b40068ad6996b
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Beginnen Sie Ihr Projekt, indem Sie ein `img`-Element mit einem `src` von `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` innerhalb des `body`-Elements hinzufügen.

# --hints--

Ihr Code sollte ein `img`-Element enthalten.

```js
assert(document.querySelectorAll('img').length === 1);
```

Das `img`-Element sollte einen `src` von `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` haben.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
