---
id: 60a3e3396c7b40068ad6996e
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Der Außenabstand ist der Bereich außerhalb der Box und kann verwendet werden, um den Abstand zwischen anderen Boxen oder Elementen zu steuern.

Hier hat das untere Element einen größeren oberen Außenabstand, wodurch es weiter nach unten auf der Seite geschoben wird.

Nachdem Sie das CSS-Boxmodell schnell überprüft haben, ist es Zeit, mit dem Rothko-Gemälde zu beginnen.

Entfernen Sie das `img`-Element.

# --hints--

Sie sollten kein `img`-Element in Ihrem Code haben.

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
