---
id: 60a3e3396c7b40068ad6996f
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Fügen Sie ein `div`-Element in das `body` ein.

Setzen Sie das `class`-Attribut gleich `canvas`.

Dies dient als Canvas für Ihre Malerei.

# --hints--

Ihr Code sollte ein `div`-Element enthalten.

```js
assert(document.querySelectorAll('div').length === 1)
```

Das `div`-Element sollte ein `class` mit dem Wert `canvas` haben.

```js
assert(document.querySelector('div').className.split(' ').includes('canvas'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
