---
id: 5dc174fcf86c76b9248c6eb2
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML Elemente haben öffnende Tags wie `<h1>` und schließende Tags wie `</h1>`. Der Text des Elementes, der angezeigt werden soll, gehört zwischen seinen öffnenden und schließenden Tags.

Ändere den Text des unten stehenden `h1`-Elements von `Hello World` zu `CatPhotoApp` und schau dir die Veränderung im Browserfenster an.

Wenn du fertig bist, klick den "Überprüfe deinen Code" (Check Your Code) Knopf und schau, ob dein Code richtig ist.

# --hints--

Der Text `CatPhotoApp` sollte im Code vorhanden sein. Überprüfe deine Rechtschreibung.

```js
assert(code.match(/catphotoapp/i));
```

Dein `h1`-Element sollte einen öffnenden Tag haben. Öffnende Tags haben folgende Syntax: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Dein `h1`-Element sollte einen schließenden Tag haben. Schließende Tags haben ein `/` nach dem `<` Zeichen.

```js
assert(code.match(/<\/h1\>/));
```

Du hast mehr als ein `h1`-Element. Entferne das überflüssige `h1`-Element.

```js
assert(document.querySelectorAll('h1').length === 1);
```

Der Text deines `h1`-Elements sollte `CatPhotoApp` sein. Du hast entweder den Text vergessen, einen Tippfehler gemacht oder der Text befindet sich nicht zwischen den öffnenden und schließenden Tags deines `h1`-Elements.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Du hast wahrscheinlich eine Browser-Erweiterung, die die Seite modifiziert. Stelle sicher, dass du alle Browser-Erweiterungen deaktivierst.

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 1);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

