---
id: 5dc17d3bf86c76b9248c6eb4
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Das `p`-Element wird verwendet, um einen Textabsatz auf den Webseiten zu erstellen. Erstelle ein `p`-Element unter deinem `h2`-Element und gib ihm den folgenden Text:

`Everyone loves cute cats online!`

# --hints--

Dein `p`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben die folgende Syntax: `<elementName>`.

```js
assert(document.querySelector('p'));
```

Dein `p`-Element sollte ein schließendes Tag haben. Abschließende Tags enthalten ein `/` direkt nach dem `<`-Zeichen.

```js
assert(code.match(/<\/p\>/));
```

Dein `p` Elementtext sollte `Everyone loves cute cats online!` lauten. Du hast entweder den Text weggelassen oder hast einen Tippfehler im Text.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

Dein `p`-Element sollte unter dem `h2`-Element stehen. Sie sind in falscher Reihenfolge.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>

--fcc-editable-region--
  </body>
</html>
```

