---
id: 5dc17d3bf86c76b9248c6eb4
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

The `p` element is used to create a paragraph of text on websites. Create a `p` element below your `h2` element and give it the following text:

`See more cat photos in our gallery.`

# --hints--

Dein `p`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben die folgende Syntax: `<elementName>`.

```js
assert(document.querySelector('p'));
```

Dein `p`-Element sollte ein schließendes Tag haben. Abschließende Tags enthalten ein `/` direkt nach dem `<`-Zeichen.

```js
assert(code.match(/<\/p\>/));
```

Your `p` element's text should be `See more cat photos in our gallery.` You have either omitted the text or have a typo.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert(extraSpacesRemoved.match(/see more cat photos in our gallery\.?$/i));
```

Dein `p`-Element sollte unter dem `h2`-Element stehen. Sie sind in falscher Reihenfolge.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>

--fcc-editable-region--
  </body>
</html>
```

