---
id: 5dc23f9bf86c76b9248c6eba
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Du kannst Bilder zu deiner Webseite mit dem `img`-Element hinzufügen. `img`-Elemente haben einen öffnenden Tag aber keinen schließenden. Ein Element ohne ein schließenden Tag wird auch <dfn>void element</dfn> genannt.

Füge ein `img`-Element unter dem `p`-Element hinzu. Zu diesem Zeitpunkt wird noch kein Bild im Browser angezeigt.

# --hints--

Dein `img`-Element sollte einen öffnenden Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert(document.querySelector('img'));
```

Dein `img`-Element sollte keinen schließenden Tag haben (`</img>`).

```js
assert(!code.match(/<\/img\>/));
```

Du solltest nur ein `img`-Element haben. Entferne die überflüssigen.

```js
assert(document.querySelectorAll('img').length === 1);
```

Dein `img`-Element sollte unter dem `p`-Element sein. Du hast die Reihenfolge vertauscht.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```

