---
id: 5dc24614f86c76b9248c6ebd
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Du kannst andere Seiten mit dem Ankerelement (`a`) verlinken.

Hier ist ein Beispiel zum Verlinken auf `https://www.freecodecamp.org`:

```html
<a href="https://www.freecodecamp.org"></a>
```

Füge ein Ankerelement, das auf `https://freecatphotoapp.com` verlinkt nach dem Paragraf hinzu. Zu diesem Zeitpunkt wird dein Link noch nicht in der Vorschau angezeigt.

# --hints--

Dein (`a`)-Ankerelement sollte ein öffnendes Tag haben. Öffnende Tags haben folgende Syntax: `<elementName>`.

```js
assert(document.querySelector('a'));
```

Dein (`a`) Ankerelement sollte ein schließendes Tag haben. Schließende Tags haben ein `/` genau nach dem `<`-Zeichen.

```js
assert(code.match(/<\/a\>/));
```

Dein Anker (`a`) Element sollte unter dem `p` Element sein. Sie sind in falscher Reihenfolge.

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('A'));
```

Dein (`a`) Ankerelement hat kein `href` Attribut. Überprüfe, ob ein Leerzeichen nach dem Namen des öffnenden Tags und/oder Leerzeichen vor allen Attributnamen vorhanden sind.

```js
assert(document.querySelector('a').hasAttribute('href'));
```

Dein (`a`) Ankerelement sollte auf `https://freecatphotoapp.com` verlinken. Du hast entweder die URL weggelassen oder einen Tippfehler gemacht.

```js
assert(
  document.querySelector('a').getAttribute('href') ===
    'https://freecatphotoapp.com'
);
```

Although you have set the anchor (`a`) element's `href` attribute to the correct link, it is recommended to always surround the value of an attribute with quotation marks.

```js
assert(
  !/\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/.test(code)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>

--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

