---
id: 5ddbd81294d8ddc1510a8e56
title: Schitt 11
challengeType: 0
dashedName: step-11
---

# --description--

Der Text des Links muss zwischen dem öffnenden und dem schließenden Tag deines Ankerelements (`a`) sein.

Hier ist ein Beispiel für einen Link mit dem Text `click here to go to freeCodeCamp.org`:

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

Add the anchor text `cat photos` to the anchor element. Dies wird der Text deines Links sein.

# --hints--

Dein (`a`)-Ankerelement sollte ein öffnendes Tag haben. Öffnende Tags haben folgende Syntax: `<elementName>`.

```js
assert(document.querySelector('a'));
```

Dein (`a`) Ankerelement sollte ein schließendes Tag haben. Schließende Tags haben ein `/` genau nach dem `<`-Zeichen.

```js
assert(code.match(/<\/a\>/));
```

Your anchor (`a`) element's text should be `cat photos`. Stelle sicher, dass der Linktext des (`a`)-Ankerelements zwischen dem öffnenden und schließenden Tag ist.

```js
assert.strictEqual(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

