---
id: 5dfa2407b521be39a3de7be1
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

Um Links in einem neuen Tab zu öffnen, kannst du das `target` Attribut im Anker (`a`) Element verwenden.

Das `target`-Attribut gibt an, wo das verknüpfte Dokument geöffnet werden soll. `target="_blank"` öffnet das verknüpfte Dokument in einem neuen Tab oder Window.

Hier ist die Basis-Syntax für ein `a`-Element mit einem `target`-Attribut:

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

Füge ein `target`-Attribut mit dem Wert `_blank` zum öffnenden Tag des (`a`)-Ankerelements hinzu, sodass der Link in einem neuen Tab geöffnet wird.

# --hints--

Dein `p`-Element sollte einen verschachtelten Anker (`a`) Element mit dem Text `cat photos` haben. Du hast es vielleicht gelöscht oder einen Tippfehler gemacht.

```js
const anchor = document.querySelectorAll('p > a');
assert.isNotEmpty(anchor);
assert.strictEqual(anchor[1]?.innerText?.toLowerCase().replace(/\s+/g, ' '), 'cat photos')
```

Dein (`a`)-Ankerelement hat kein `target`-Attribut. Überprüfe, ob nach dem Namen des öffnenden Tags ein Leerzeichen steht und/oder ob vor allen Attributnamen Leerzeichen stehen.

```js
assert.isTrue(document.querySelectorAll('a')[1]?.hasAttribute('target'));
```

Der Wert des `target`-Attributs sollte `_blank` sein. Du hast entweder den Wert weggelassen oder einen Tippfehler gemacht. Denke daran, dass Attributwerte zwischen Anführungszeichen geschrieben werden sollten.

```js
assert.strictEqual(document.querySelectorAll('a')[1]?.getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

