---
id: 5dfb655eeacea3f48c6300b3
title: Schritt 25
challengeType: 0
dashedName: step-25
---

# --description--

Das `figure`-Element repräsentiert eigenständige Inhalte und ermöglicht dir, ein Bild mit einer Beschriftung zu verbinden.

Bette das Bild, das du gerade hinzugefügt hast, in ein `figure`-Element ein.

# --hints--

Dein `figure`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben die folgende Syntax: `<elementName>`.

```js
assert(document.querySelector('figure'));
```

Dein `figure` sollte ein schließendes Tag haben. Abschließende Tags enthalten ein `/` direkt nach dem `<`-Zeichen.

```js
assert(code.match(/<\/figure\>/));
```

Ein `figure`-Element sollte genau über dem schließenden Tag des zweiten `section`-Elements stehen.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.lastElementChild?.nodeName, 'FIGURE');
```

Das Lasagne-`img`-Element sollte innerhalb des `figure`-Elements eingebettet sein.

```js
assert.strictEqual(document.querySelector('figure > img')?.src.toLowerCase(), 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--

        <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

