---
id: 5dfb6a35eacea3f48c6300b4
title: Schritt 26
challengeType: 0
dashedName: step-26
---

# --description--

Ein Abbildungsbeschriftungselement (`figcaption`) wird verwendet, um eine Beschriftung hinzuzufügen, die das im `figure` enthaltene Bild beschreibt.

Hier ist ein Beispiel für ein `figcaption`-Element mit der Beschriftung von `A cute cat`:

```html
<figure>
  <img src="image.jpg" alt="A description of the image">
  <figcaption>A cute cat</figcaption>
</figure>
```

Füge nach dem im `figure`-Element eingebetteten Bild ein `figcaption`-Element mit dem folgenden Text hinzu:

`Cats love lasagna.`

# --hints--

Dein `figcaption` Element sollte ein öffnendes Tag haben. Öffnende Tags haben die folgende Syntax: `<elementName>`.

```js
assert(document.querySelector('figcaption'));
```

Dein `figcaption` Element sollte ein schließendes Tag haben. Abschließende Tags haben ein `/` genau nach dem `<`-Zeichen.

```js
assert(code.match(/<\/figcaption\>/));
```

Das `figcaption` Element sollte im `figure` Element eingebettet sein.

```js
assert(
  document.querySelector('figure > figcaption') &&
    document.querySelector('figure > figcaption')
);
```

Das `img` Lasagna-Element sollte im `figure` Element eingebettet sein.

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

Das `figcaption`-Element, das im `figure`-Element eingebettet wurde, sollte über dem `img`-Element sein. Sie sind in falscher Reihenfolge.

```js
assert(
  document.querySelector('figcaption').previousElementSibling.nodeName === 'IMG'
);
```

Der Text deines `figcaption`-Elements sollte `Cats love lasagna.` sein. Du hast ihn entweder vergessen oder einen Tippfehler gemacht.

```js
assert(
  document.querySelector('figcaption').innerText.match(/Cats love lasagna.?$/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

