---
id: 5ef9b03c81a63668521804dd
title: Schritt 47
challengeType: 0
dashedName: step-47
---

# --description--

`label`-Elemente werden verwendet, um den Text für ein `input`-Element mit dem `input`-Element selbst zu verknüpfen (speziell für unterstützende Technologien wie Screenreader).

Hier ist ein Beispiel für ein `label`-Element mit einem `radio`-Button:

```html
<label><input type="radio"> cat</label>
```

Wenn du nun in der Vorschau auf das Wort `"cat"` klickst, wird die `radio`-Schaltfläche auch ausgewählt.

Bette deinen `radio`-Button in einem `label`-Element ein.

# --hints--

Stelle sicher, dass der Radio-Button noch vorhanden ist.

```js
assert(document.querySelector('input[type="radio"]'));
```

Der Text `Indoor` sollte sich direkt rechts neben deinem `radio`-Button befinden. Stelle sicher, dass zwischen dem Element und dem Text ein Leerzeichen ist. Du hast entweder den Text weggelassen oder einen Tippfehler gemacht.

```js
const radioInputElem = document.querySelector('input');
assert.match(
  radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '),
  / Indoor/i
);
```

Dein `label`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben folgende Syntax: `<elementName>`.

```js
assert(document.querySelector('label'));
```

Dein `label`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` nach dem `<` Zeichen.

```js
assert(code.match(/<\/label\>/));
```

Dein Radio-Button und sein Text sollten sich zwischen dem öffnenden und schließenden Tag des `label`-Elements befinden.

```js
const labelChildNodes = [...document.querySelector('form > label').childNodes];
assert(
  labelChildNodes.filter((childNode) => childNode.nodeName === 'INPUT').length
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="radio"> Indoor
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

