---
id: 5ef9b03c81a63668521804df
title: Step 48
challengeType: 0
dashedName: step-48
---

# --description--

Das `id`-Attribut wird verwendet, um bestimmte HTML-Elemente zu identifizieren. Jeder Wert des `id`-Attributs darf auf der gesamten Seite nur einmal vorhanden sein, muss sich also von allen anderen `id`-Werten unterscheiden.

Hier ist ein Beispiel für ein `input`-Element mit einem `id`-Attribut:

```html
<input id="email">
```

Füge ein `id`-Attribut mit dem Wert `indoor` zum Radio-Button hinzu. Wenn Elemente mehrere Attribute haben, spielt die Reihenfolge der Attribute keine Rolle.

# --hints--

Dein Radio-Button sollte immer noch zwischen dem öffnenden und schließenden Tag des `label`-Elements stehen.

```js
const labelChildNodes = [...document.querySelectorAll('form > label')?.[0]?.childNodes];
assert(
  labelChildNodes.filter((childNode) => childNode.nodeName === 'INPUT').length
);
```

Dein Radio-Button sollte ein `id`-Attribut haben. Überprüfe, dass ein Leerzeichen nach dem Namen des öffnendes Tags und/oder Leerzeichen vor allen Attributnamen vorhanden sind.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('id'));
```

Dein Radio-Button sollte ein `type`-Attribut haben. Überprüfe, ob nach dem Namen des öffnenden Tags ein Leerzeichen steht und/oder ob vor allen Attributnamen Leerzeichen stehen.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

Dein Radio-Button sollte ein `id`-Attribut mit dem Wert `indoor` haben. Du hast entweder den Wert weggelassen oder einen Tippfehler gemacht. Denke daran, dass Attributwerte zwischen Anführungszeichen geschrieben werden sollten.

```js
assert.match(document.querySelector('input')?.id, /^indoor$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input type="radio"> Indoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

