---
id: 5ef9b03c81a63668521804dc
title: Schritt 46
challengeType: 0
dashedName: step-46
---

# --description--

Du kannst Radiobuttons für Fragen verwenden, bei denen du nur eine Antwort aus mehreren Optionen möchtest.

Hier ist ein Beispiel für ein Radio-Button mit der Option von `cat`:

```html
<input type="radio"> cat
```

Denk daran, dass ein `input`-Elemnet ein leeres Element ist.

Füge vor der Text-Input einen Radio-Button hinzu, der diese Option enthält:

`Indoor`

# --hints--

Du solltest für deinen Radio-Button ein `input`-Element erstellen. Überprüfe die Syntax.

```js
assert(document.querySelectorAll('form > input').length >= 2);
```

Dein `input`-Element sollte einen öffnenden Tag, aber keinen schließenden Tag haben.

```js
assert(document.querySelector('form > input') && !code.match(/<\/input\>/g));
```

Du solltest nur ein `input`-Element für deinen Radio-Button hinzugefügt haben. Entferne die überflüssigen.

```js
assert(document.querySelectorAll('form > input'), 2);
```

Dein neues `input`-Element sollte über dem vorhandenen `input` sein und ein `type` Attribut auf `text` gesetzt haben. Sie sind in falscher Reihenfolge.

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

Dein neues `input`-Element hat kein `type`-Attribut. Überprüfe, ob ein Leerzeichen nach dem Namen des öffnenden Tags vorhanden ist.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

Dein neues `input`-Element sollte nur ein `type`-Attribut haben. Entferne die überflüssigen.

```js
assert(document.querySelector('input')
  ?.getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

Dein neues `input`-Element sollte ein `type`-Attribut mit dem Wert `radio` haben. Du hast entweder den Wert weggelassen oder dir ist ein Tippfehler unterlaufen. Denke daran, dass Attributwerte von Anführungszeichen umgeben sein sollten.

```js
assert.match(
  document.querySelector('input')?.getAttribute('type'),
  /^radio$/i
);
```

Obwohl du das `type`-Attribut des neuen `input`-Elements auf `radio` gesetzt hast, wird empfohlen, den Wert eines Attributs immer zwischen Anführungszeichen zu schreiben.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

Der `Indoor`-Text des `radio`-Buttons sollte danach und nicht davor stehen.

```js
const radioInputElem = document.querySelector('input');
assert.notMatch(radioInputElem?.previousSibling?.nodeValue, /Indoor/i);
```

Der Text `Indoor` sollte sich direkt rechts neben deinem `radio`-Button befinden. Du hast entweder den Text vergessen oder hast einen Tippfehler gemacht.

```js
const radioInputElem = document.querySelector('input');
assert.match(radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '), /\s*Indoor/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--

          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

