---
id: 5ef9b03c81a63668521804d0
title: Schritt 27
challengeType: 0
dashedName: step-27
---

# --description--

Um die Betonung auf ein bestimmtes Wort oder eine bestimmte Phrase zu legen, kannst Du das Element `em` verwenden.

Hebe das Wort `love` im `figcaption`-Element hervor, indem du es in ein betonendes `em`-Element, hinein setzt.

# --hints--

Dein `em`-Betonungselement sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert(document.querySelector('em'));
```

Dein `em`-Betonungselement sollte ein schließendes Tag haben. Abschließende Tags enthalten ein `/` direkt nach dem `<`-Zeichen.

```js
assert(code.match(/<\/em\>/));
```

Du hast entweder das `figcaption` Element gelöscht oder es fehlt ein öffnendes oder schließendes Tag.

```js
assert(document.querySelector('figcaption') && code.match(/<\/figcaption\>/));
```

Dein `em`-Betonungselement sollte den Text `love` umfassen. Du hast entweder den Text weggelassen oder einen Tippfehler gemacht.

```js
assert(
  document.querySelector('figcaption > em').innerText.toLowerCase() === 'love'
);
```

Der `figcaption`-Text sollte `Cats love lasagna` lauten. Prüfe, ob Tippfehler und die notwendigen Leerzeichen um das öffnende und abschließende Tag des `em`-Elements vorhanden sind.

```js
assert(
  document
    .querySelector('figcaption')
    .innerText.replace(/\s+/gi, ' ')
    .match(/cats love lasagna\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

