---
id: 5ef9b03c81a63668521804db
title: Step 43
challengeType: 0
dashedName: step-43
---

# --description--

Um zu verhindern, dass ein Benutzer dein Formular absendet, wenn erforderliche Informationen fehlen, musst du das `required`-Attribut zu einem `input`-Element hinzufügen. Es ist nicht notwendig, dem `required`-Attribut einen Wert zu geben. Füge stattdessen einfach das Wort `required` zum `input`-Element hinzu und stelle sicher, dass zwischen ihm und anderen Attributen Leerzeichen vorhanden sind.

# --hints--

Du hast entweder dein `input`-Element gelöscht oder es hat eine ungültige Syntax. Alle Attributwerte sollten zwischen Anführungszeichen stehen.

```js
assert(document.querySelectorAll('input').length);
```

Deine `form` sollte nur das `input`-Element enthalten. Entferne innerhalb des `form`-Elements alle zusätzlichen HTML-Elemente oder Text.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Dein `input`-Element sollte das `required`-Attribut besitzen. Denke daran, dass du nur das Wort `required` in das Tag des `input`-Elements hinzufügst.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('required'));
```

Dem `required`-Attribut sollte kein Wert gegeben werden.

```js
assert.strictEqual(document.querySelector('input')?.getAttribute('required'), '');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

