---
id: 5efc4f528d6a74d05e68af74
title: Schritt 58
challengeType: 0
dashedName: step-58
---

# --description--

Es gibt eine andere Möglichkeit, den Text eines `input`-Elements mit dem Element selbst zu verknüpfen. Du kannst den Text in einem `label`-Element einbetten und dann das `for`-Attribut mit demselben Wert, wie die `id` des `input`-Elements hinzufügen.

Gegeben ist ein `input`-Element wie unten:

```html
<input id="breakfast" type="radio" name="meal" value="breakfast">
```

Ein Beispiel für ein `label` Element das mit diesem `input`-Element in Verbindung gebracht wird ist:

```html
<label for="breakfast">Breakfast</label>
```

Verknüpfe den Text `Loving` mit der Checkbox, indem du nur den Text `Loving` in einem `label`-Element verschachtelst und ihm ein entsprechendes `for`-Attribut gibst.

# --hints--

Du solltest sicherstellen, dass das Kontrollkästchen noch vorhanden ist und Du keine Änderungen daran vorgenommen hast. Es sollte so bleiben, wie es im Anfangscode war.

```js
const checkboxElementRegex = /\s*<input\s+id\s*=\s*"loving"\s*type\s*=\s*"checkbox"\s*>/;
assert.match(code, checkboxElementRegex);
```

Dein Kontrollkästchen sollte ein `id`-Attribut mit dem Wert `loving` haben.  Es wurde `--fcc-expected--` erwartet, allerdings wurde `--fcc-actual--` gefunden.

```js
const checkboxElementId = document.querySelector('input[type="checkbox"]')?.getAttribute('id');
assert.strictEqual(checkboxElementId, 'loving');
```

Du solltest nach dem Kontrollkästchen ein neues `label`-Element hinzufügen. Stelle sicher, dass es sowohl ein öffnendes als auch schließendes Tag enthält.

```js
assert.lengthOf(document.querySelectorAll('label'), 3);
assert.lengthOf(code.match(/<\/label\>/g), 3);
```

Dein `label` Element sollte nach dem Kontrollkästchen sein.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElement = checkboxElement?.nextElementSibling?.nodeName;
assert.strictEqual(checkboxSiblingElement, 'LABEL');
```

Dein `label`-Element sollte den Text `Loving` haben.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementText = checkboxElement?.nextElementSibling?.innerText;

assert.strictEqual(checkboxSiblingElementText, 'Loving');
```

Der Text `Loving` sollte innerhalb des neuen `label`-Elements sein. Stelle sicher, dass Du den Text nicht außerhalb des `label` Elements vervielfältigt hast.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingNode = checkboxElement?.nextSibling?.nodeValue ?? '';
const labelElement = checkboxElement?.nextElementSibling;
const labelSiblingNode = labelElement?.nextSibling?.nodeValue ?? '';

assert.notInclude(checkboxSiblingNode, 'Loving');
assert.notInclude(labelSiblingNode, 'Loving');
```

Dein `label` Element hat kein `for` Attribut. Stell sicher, dass ein Leerzeichen zwischen dem sich öffnenden `label` Tag -Namen und dem `for` Attribut vorhanden ist.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementAttr = checkboxElement?.nextElementSibling?.getAttribute('for');

assert.isNotNull(checkboxSiblingElementAttr);
```

Dein `label`-Element sollte ein `for`-Attribut mit dem Wert `loving` enthalten.  Es wurde `--fcc-expected--` erwartet, allerdings hast du `--fcc-actual--` geschrieben.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementAttr = checkboxElement?.nextElementSibling?.getAttribute('for');

assert.strictEqual(checkboxSiblingElementAttr, 'loving');
```

Zwischen deiner Checkbox und deinem neuen `label`-Element sollte ein Leerzeichen sein.

```js
assert.match(code, />\s+<label\s+for\s*=\s*('|")loving/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> Loving
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

