---
id: 5f07be6ef7412fbad0c5626b
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

Bevor du neuen Inhalt hinzufügst, solltest du ein `section`-Element verwenden, um die Katzenfotos vom zukünftigen Inhalt zu trennen.

Das `section`-Element wird verwendet, um Abschnitte in einem Dokument zu definieren, wie Kapitel, Kopfzeilen, Fußzeilen oder andere Abschnitte des Dokuments. Das `section`-Element ist ein semantisches Element, das bei SEO und Barrierefreiheit hilft.

```html
<section>
  <h2>Section Title</h2>
  <p>Section content...</p>
</section>
```

Nimm dein `h2`, den Kommentar, das `p` und das Anker (`a`) Element und bette sie in ein `section`-Element ein.

# --hints--

Dein `section`Element sollte ein öffnendes Tag haben. Öffnende Tags haben die folgende Syntax: `<elementName>`.

```js
assert(document.querySelector('section'));
```

Dein `section`-Element sollte ein abschließendes Tag haben. Schließende Tags haben ein `/` genau nach dem `<`-Zeichen.

```js
assert(code.match(/<\/section\s*>/i));
```

Das ganze `section`-Element sollte sich zwischen dem öffnenden und schließenden Tag des `main`-Elements befinden.

```js
assert(document.querySelector('section').parentNode.nodeName === 'MAIN');
```

The existing `h2`, comment, two `p` elements, and anchor (`a`) element should be between the opening and closing tags of the `section` element.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert.lengthOf(foundElements, 4);
```

Das `h1`-Element sollte nicht im `section`-Element verschachtelt werden.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>

      <h2>Cat Photos</h2>
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

    </main>
--fcc-editable-region--
  </body>
</html>
```

