---
id: 671141f948cbab359e74cc93
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Füge `p` Tags hinzu, um `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.` in einen Absatz (englisch paragraph) umzuwandeln.

# --hints--

Du solltest ein `p` Element um `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.` herum hinzufügen.

```js
const P = document.querySelectorAll('p')[1];
assert.exists(P);
assert.strictEqual(P.innerText, "See more cat photos in our gallery.");
```

Du solltest immer noch das Ankerelement im neuen Absatz haben.

```js
const P = document.querySelectorAll('p')[1];
assert.exists(P.querySelector('a'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
