---
id: 671141feba228a35cefba82d
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

Turn the existing text `cute cats` into an anchor element that links to:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`

# --hints--

Es sollte ein neues Ankerelement im ersten `p` Element geben.

```js
assert.exists(document.querySelector('p')?.querySelector('a'));
```

Der Ankertext sollte `cute cats` sein.

```js
assert.strictEqual(document.querySelector('a')?.innerText, "cute cats");
```

Das `href` Attribut des neuen Ankerelements sollte `https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg` sein.

```js
assert.strictEqual(document.querySelector('a')?.href, "https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg");
```

Es sollte nur einmal `cute cats` in deinem Code geben.

```js
assert.strictEqual(code.indexOf("cute cats"),code.lastIndexOf("cute cats")); 
```

Der Text des `p` Elements sollte immer noch `Everyone loves cute cats online!`.

```js
assert.strictEqual(document.querySelector('p')?.innerText, "Everyone loves cute cats online!");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
