---
id: 671141d8e32fe934c26fa1be
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Füge die Wörter `See more` vor dem Ankerelement und `in our gallery` nach dem Ankerelement ein.

# --hints--

Du solltest das vorhandene Ankerelement verwenden, kein neues erstellen.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

Der Text innerhalb des Ankerelements darf sich nicht ändern, es muss `cat photos` sein.

```js
assert.strictEqual(document.querySelector('a').innerText, "cat photos");
```

Du solltest `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery` in deinem Code haben.

```js
assert.match(code, /See more <a href="https:\/\/freecatphotoapp\.com">cat photos<\/a> in our gallery/)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
