---
id: 671144cdcc01d73f7dd79dc9
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Nun, da Du den Link hinzugefügt hast, kannst du den Kommentar entfernen.

# --hints--

Du solltest keine Kommentare in deinem Code haben.

```js
assert.notMatch(code, /<!--/);
assert.notMatch(code, /-->/);
```

Der Text `TODO: Add link to cat photos` sollte nicht mehr vorhanden sein.

```js
assert.notMatch(code, /TODO:?/i)
assert.notMatch(code, /Add\s*link\s*to\s*cat\s*photos/i)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
--fcc-editable-region--
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
