---
id: 7cf9b03d81a65668421804c3
title: Schritt 41
challengeType: 0
dashedName: step-41
---

# --description--

Damit der Server, an den deine Daten geschickt werden (`action`-Attribut), auf die Daten eines Formulars zugreifen kann, musst du dem Textfeld ein `name`-Attribut mit einem Wert geben, der die übermittelten Daten benennt.

Hier ist ein Beispiel für ein `input`-Element mit einem `name`-Attribut:

```html
<input type="text" name="name">
```

Füge das `name`-Attribut mit dem Wert `catphotourl` zu deinem Textfeld hinzu.

# --hints--

Du hast entweder dein `input`-Element gelöscht oder es hat eine ungültige Syntax. Alle Attributwerte sollten zwischen Anführungszeichen stehen.

```js
assert(document.querySelectorAll('input').length);
```

Deine `form` sollte nur das `input`-Element enthalten. Entferne innerhalb des `form`-Elements alle zusätzlichen HTML-Elemente oder Text.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
  document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Dein `input`-Element hat kein `name`-Attribut. Überprüfe, ob nach dem Namen des öffnenden Tags ein Leerzeichen steht und/oder ob vor allen Attributnamen Leerzeichen stehen.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('name'));
```

Dein `input`-Element sollte ein `name`-Attribut mit `catphotourl` als Wert haben.  Du hast entweder den Wert weggelassen oder einen Tippfehler gemacht. Denke daran, dass Attributwerte zwischen Anführungszeichen geschrieben werden sollten.

```js
assert.match(
  document.querySelector('input')?.getAttribute('name'),
  /^catphotourl$/i
);
```

Obwohl du das `name`-Attribut des `input`-Elements auf `catphotourl` gesetzt hast, wird es empfohlen, die Werte eines Attributs immer zwischen Anführungszeichen zu schreiben.

```js
assert(!/\<\s*input\s+.*\s*=\s*catphotourl/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

