---
id: 69a92f0b9cc04bb0d5327bb5
title: Реалізувати стек
challengeType: 26
dashedName: implement-a-stack
---

# --description--

У цій лабораторній роботі ви реалізуєте структуру даних стек за допомогою функцій. Стек — це структура даних типу «останній зайшов — перший вийшов» (LIFO), де елементи додаються та видаляються з вершини.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історія користувача:**

1. У вас має бути функція `initStack`, яка повертає об’єкт із властивістю `collection`, встановленою в порожній масив. Останній елемент `collection` представляє верхівку стеку.
1. У вас має бути функція `push`, яка приймає об’єкт стеку та елемент як аргументи і додає цей елемент на верхівку стеку, розміщуючи його в кінці `collection` стеку.
1. У вас має бути функція `pop`, яка приймає об’єкт стеку як аргумент, видаляє та повертає верхній елемент стеку або `undefined`, якщо такого немає.
1. У вас має бути функція `peek`, яка приймає об’єкт стеку як аргумент і повертає верхній елемент стеку без його видалення, або `undefined`, якщо такого немає.
1. У вас має бути функція `isEmpty`, яка приймає об’єкт стеку як аргумент і повертає `true`, якщо стек не містить елементів, і `false` в іншому випадку.
1. У вас має бути функція `clear`, яка приймає об’єкт стеку як аргумент і видаляє всі елементи з `collection` стеку.

**Примітка:** Більшість тестів залежать від функцій `initStack` і `push`. Спочатку реалізуйте їх, оскільки тести для `pop`, `peek`, `isEmpty` і `clear` вимагають додавання елементів у стек.

# --hints--

Ви повинні мати функцію `push`.

```js
assert.isFunction(push);
```

Функція `push` має додавати елемент на верхівку стеку.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

Ви повинні мати функцію `pop`.

```js
assert.isFunction(pop);
```

Функція `pop` має видаляти та повертати верхній елемент стеку.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
assert.strictEqual(pop(stack), 'third');
assert.deepEqual(stack.collection, ['first', 'second']);
assert.strictEqual(pop(stack), 'second');
assert.deepEqual(stack.collection, ['first']);
```

Функція `pop` має повертати `undefined`, якщо стек порожній.

```js
const stack = initStack();
assert.isUndefined(pop(stack));
```

Функція `pop` має правильно повертати хибні значення.

```js
const stack = initStack();
push(stack, '');
assert.strictEqual(pop(stack), '');
assert.lengthOf(stack.collection, 0);
```

Ви повинні мати функцію `peek`.

```js
assert.isFunction(peek);
```

Функція `peek` має повертати верхній елемент стеку без його видалення.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.strictEqual(peek(stack), 'second');
assert.strictEqual(peek(stack), 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

Функція `peek` має повертати `undefined`, якщо стек порожній.

```js
const stack = initStack();
assert.isUndefined(peek(stack));
```

Функція `peek` має правильно повертати хибні значення.

```js
const stack = initStack();
push(stack, 0);
assert.strictEqual(peek(stack), 0);
```

Ви повинні мати функцію `isEmpty`.

```js
assert.isFunction(isEmpty);
```

Функція `isEmpty` має повертати `true` для порожнього стеку.

```js
const stack = initStack();
assert.isTrue(isEmpty(stack));
```

Функція `isEmpty` має повертати `false` для непорожнього стеку.

```js
const stack = initStack();
push(stack, 'element');
assert.isFalse(isEmpty(stack));
```

Функція `isEmpty` має повертати `false`, коли верхній елемент є хибним значенням.

```js
const stack = initStack();
push(stack, false);
assert.isFalse(isEmpty(stack));
```

Ви повинні мати `clear` функцію.

```js
assert.isFunction(clear);
```

Функція `clear` має видаляти всі елементи зі стеку.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
clear(stack);
assert.lengthOf(stack.collection, 0);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function initStack() {
  return {
    collection: []
  };
}

function print(stack) {
  console.log(stack.collection);
}

function push(stack, element) {
  stack.collection.push(element);
}

function pop(stack) {
  return stack.collection.pop();
}

function peek(stack) {
  return stack.collection[stack.collection.length - 1];
}

function isEmpty(stack) {
  return stack.collection.length === 0;
}

function clear(stack) {
  stack.collection = [];
}
```
