---
id: 6852dd8486c7ee0171322950
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Você precisará de uma variável para armazenar o código secreto enquanto o decodifica, então remova `pass` e crie uma variável chamada `secret_code` dentro da função e atribua a ela uma string vazia.

# --hints--

Você deve ter uma variável `secret_code` dentro da função `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_code')
    `)
})
```

A variável `secret_code` deve ter um valor de uma string vazia.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_code').is_equivalent('secret_code = ""')
    `)
})
```

Você não deve mais ter `pass` no seu código.

```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').has_pass()
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    pass
--fcc-editable-region--
```
