---
id: 6852db755c9c1274d75aae9a
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Nesta oficina, você vai criar um extrator de PIN. Os dígitos do PIN estão escondidos em cada linha de um poema.

Comece criando uma função chamada `pin_extractor` com um parâmetro `poem`. Uma função sem corpo gera um erro, então adicione `pass` dentro da função como um marcador de posição temporário para torná-la válida.

# --hints--

Você deve ter uma função chamada `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('pin_extractor')
    `)
})
```

A função `pin_extractor` deve ter um parâmetro chamado `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poem')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
