---
id: 68531fbb934d735415053274
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Agora que você adicionou uma chamada de função, pode ver qualquer erro criado dentro da função no terminal e sempre pode usar `print` para ver o valor das coisas.

Há um dígito do pin escondido em cada linha, então dentro da função use o método `split` para dividir a string em uma lista de linhas. Divida as linhas no caractere de nova linha (`\n`) e atribua a lista resultante a uma variável chamada `lines`.

# --hints--

Você deve ter uma variável chamada `lines` dentro da sua função `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('lines')
    `)
})
```

A variável `lines` deve ter um valor de `poem.split('\n')`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('lines').is_equivalent('lines = poem.split("\\\\n")')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
```
