---
id: 68530231c05b93324e071104
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Enquanto você trabalha na função, seria útil ter uma chamada de função para fins de depuração, então crie uma variável `poem`, fora da função, que contenha este poema:

```md
Stars and the moon
shine in the sky
white and bright
until the end of the night
```

Você pode usar uma string de múltiplas linhas para isso.

E então, chame a função com a variável `poem` como argumento.

# --hints--

Você deve ter uma variável global chamada `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_variable('poem')
    `)
})
```

A variável `poem` deve ter um valor igual a `Stars and the moon\nshine in the sky\nwhite and bright\nuntil the end of the night`.

```js
({
    test: () => runPython(`
    pm = """Stars and the moon
    shine in the sky
    white and bright
    until the end of the night"""
    assert _Node(_code).find_variable('poem').is_equivalent(f'poem = """{pm}"""')
    `)
})
```

Você deve chamar a função `pin_extractor` com `poem` como argumento.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('pin_extractor(poem)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''

--fcc-editable-region--

--fcc-editable-region--
```
