---
id: 685578aae65e4106ad2b4300
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Como você aprendeu em uma lição anterior, a função `enumerate` permite acompanhar o índice enquanto percorre um iterável:

```py
fruits = ['apple', 'plum', 'bananas'] 

for index, item in enumerate(fruits):  
    print(index, item)  

# 0 apple  
# 1 plum
# 2 bananas
```

O dígito n-ésimo do código secreto vem da palavra n-ésima da linha n-ésima, então você precisa saber qual é o número de cada linha.

Altere o loop para que ele itere em `enumerate(lines)` e adicione outra variável de loop antes de `line` chamada `line_index`.

Além disso, atualize a chamada `print` para `print(line_index, line)`.

# --hints--

O loop deve iterar sobre `enumerate(lines)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('enumerate(lines)')
    `)
})
```

O loop deve ter `line_index, line` como variável do loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line_index, line')
    `)
})
```

Você deve ter `print(line_index, line)` dentro do loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line_index, line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        print(line)
        words = line.split()
        print(words)

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
--fcc-editable-region--

```
