---
id: 685535bf527b2c77236b3a9a
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

O dígito n-ésimo do pin está oculto como o comprimento da n-ésima palavra na n-ésima linha.

Para encontrar o comprimento da enésima palavra, o próximo passo é separar a linha do poema em uma lista de palavras: dentro do loop, crie uma variável `words` e atribua a ela o valor de `line` dividido em palavras usando o método `split`.

Então, adicione uma chamada `print` com `words` como seu argumento.

# --hints--

Você deve ter uma variável `words` dentro do loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_variable('words')
    `)
})
```

A variável `words` deve ter um valor de `line.split()`.

```js
({
    test: () => runPython(`
    variable = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_variable('words')
    assert variable.is_equivalent('words = line.split()') or variable.is_equivalent('words = line.split(" ")')
    `)
})
```

Você deve usar `print(words)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        print(line)
--fcc-editable-region--
        
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
