---
id: 68553423d775cc60134ac8bb
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Você precisa trabalhar em cada linha independentemente: crie um loop sobre `lines` que use `line` como variável do loop. Dentro do loop, imprima `line`.

# --hints--

Você deve ter um loop `for` na sua função `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0]
    `)
})
```

O loop deve iterar sobre `lines`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('lines')
    `)
})
```

O loop deve ter `line` como variável do loop (`for line in lines`).

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line')
    `)
})
```

Você deve ter `print(line)` dentro do loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
