---
id: 68559f3561d7df287586142e
title: Passo 16
challengeType: 20
dashedName: step-16
---

# --description--

Agora o extrator de pinos funciona, mas poderia ser mais eficiente!Ele poderia extrair o pino de muitos poemas ao mesmo tempo!

Adicione mais dois poemas no escopo global, `poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'` e `poem3 = 'There\nonce\nwas\na\ndragon'`, e por enquanto, comente o `print` com a chamada para `pin_extractor`.


# --hints--

As variáveis `poem2` e `poem3` devem existir e ter o valor correto.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('poem2')
        assert _Node(_code).has_variable('poem3')
        assert _Node(_code).find_variable('poem2').is_equivalent("poem2 = 'The grass is green\\\\nhere and there\\\\nhoping for rain\\\\nbefore it turns yellow'")
        assert _Node(_code).find_variable('poem3').is_equivalent("poem3 = 'There\\\\nonce\\\\nwas\\\\na\\\\ndragon'")
    `)
})
```

Você deve comentar a chamada `print(pin_extractor(poem))`.

```js
({
    test: () => runPython(`
        assert not _Node(_code).has_call('print(pin_extractor(poem))')`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
        else:
            secret_code += '0'
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

print(pin_extractor(poem))
--fcc-editable-region--

```
