---
id: 6855918e87c33c6979d21d3f
title: Passo 14
challengeType: 20
dashedName: step-14
---

# --description--

Agora a função está fazendo um bom trabalho em extrair o código do poema, mas há um caso limite a considerar: a linha do poema pode ser mais curta do que o esperado e isso faria a função gerar um erro (você pode tentar isso removendo a palavra `bright` do poema).

Coloque a linha `secret_code += str(len(words[line_index]))` em uma declaração `if` que verifica se há palavras suficientes na lista `words`.

# --hints--

Dentro do loop você deve ter uma declaração `if`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0]
    `)
})
```

A condição da instrução if deve ser `len(words) > line_index`.


```js
({
    test: () => runPython(`
    _cond = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
    assert _cond.is_equivalent('len(words) > line_index') or _cond.is_equivalent('line_index < len(words)')
    `)
})
```

Dentro da declaração `if` você deve ter `secret_code += str(len(words[line_index]))`.


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt('secret_code += str(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
--fcc-editable-region--
        secret_code += str(len(words[line_index]))
--fcc-editable-region--
    return secret_code

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

print(pin_extractor(poem))

```
