---
id: 685926db99041c5e619ead41
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

A função `len` retorna o comprimento de uma string.Atualize a segunda chamada de `print` para imprimir o comprimento dessa palavra.

# --hints--

Você deve ter `print(len(words[line_index]))` no final do loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
        print(words[line_index])


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
--fcc-editable-region--

```
