---
id: 6859a037b3d6721c7fbb8055
title: Passo 18
challengeType: 20
dashedName: step-18
---

# --description--

Antes do loop, crie uma nova variável `secret_codes` que tenha um valor inicial de uma lista vazia.Substitua a linha `return secret_code` por uma linha que adicione `secret_code` à lista `secret_codes`.

# --hints--

Você deve ter uma variável `secret_codes` dentro da função `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_codes')
    `)
})
```

A variável `secret_codes` deve ter um valor de uma lista vazia.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_codes').is_equivalent('secret_codes = []')
    `)
})
```

No final do loop você deve ter `secret_codes.append(secret_code)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('secret_codes.append(secret_code)')
    `)
})
```

Você não deve retornar `secret_code`.


```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_return('secret_code')
    `)
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poems):
    
    for poem in poems:
        secret_code = ''
        lines = poem.split('\n')
        for line_index, line in enumerate(lines):
            words = line.split()
            if len(words) > line_index:
                secret_code += str(len(words[line_index]))
            else:
                secret_code += '0'
        return secret_code
        
poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'
poem3 = 'There\nonce\nwas\na\ndragon'

# print(pin_extractor(poem))
--fcc-editable-region--

```
