---
id: 68592730aaf34b61ae279168
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

O comprimento retornado de `len` é um inteiro, você precisa convertê-lo para uma string usando a função `str`.

Atualize a chamada `print` para imprimir o comprimento como uma string.

# --hints--

Você deve ter `print(str(len(words[line_index])))` no final do loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(str(len(words[line_index])))')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
        print(len(words[line_index]))


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
--fcc-editable-region--

```
