---
id: 615f3949f58e12577dcefb00
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Guter Einsatz von Weißraum kann den Fokus auf die wichtigen Elemente Ihrer Seite lenken und die Augen Ihrer Nutzer durch den Text führen.

Geben Sie Ihrem `.label`-Selektor eine `width`-Eigenschaft mit dem Wert `270px`.

# --hints--

Ihr `.label`-Selektor sollte eine `width`-Eigenschaft mit dem Wert `270px` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.label')?.width === '270px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

--fcc-editable-region--
.label {
  border: 2px solid black;
}
--fcc-editable-region--
```
