---
id: 615f2abbe7d18d49a1e0e1c8
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Wir haben Ihnen einen grundlegenden HTML-Boilerplate-Code bereitgestellt.

Erstellen Sie ein `h1`-Element innerhalb Ihres `body`-Elements und geben Sie ihm den Text `Nutrition Facts`.

# --hints--

Sie sollten ein neues `h1`-Element hinzufügen.

```js
assert.exists(document.querySelector('h1'));
```

Ihr `h1`-Element sollte sich innerhalb Ihres `body`-Elements befinden.

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

Ihr `h1`-Element sollte den Text `Nutrition Facts` enthalten.

```js
assert(document.querySelector('h1')?.innerText === 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>

</body>
</html>
--fcc-editable-region--
```

```css

```
