---
id: 615f3cafd794015aa9547a6d
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Denken Sie daran, dass die Verwendung von `h1`, `h2` und ähnlichen Tags die semantische Struktur Ihres HTML bestimmt. Sie können jedoch das CSS dieser Elemente anpassen, um den visuellen Fluss und die Hierarchie zu steuern.

Erstellen Sie eine `h1`-Regel und setzen Sie die `font-weight`-Eigenschaft auf `800`. Dadurch wird Ihr `h1`-Text fetter dargestellt.

# --hints--

Sie sollten einen `h1`-Selektor erstellen.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

Ihr `h1`-Selektor sollte eine `font-weight`-Eigenschaft mit dem Wert `800` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === '800');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--

--fcc-editable-region--
```
