---
id: 615f3ed16592445e57941aec
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Linien können helfen, wichtigen Inhalt zu trennen und zu gruppieren, besonders wenn der Platz begrenzt ist.

Erstellen Sie ein `div`-Element unterhalb Ihres `h1`-Elements und geben Sie ihm ein `class`-Attribut mit dem Wert `divider`.

# --hints--

Sie sollten ein neues `div` erstellen.

```js
assert(document.querySelectorAll('div')?.length === 2);
```

Ihr neues `div` sollte das `class`-Attribut mit dem Wert `divider` haben.

```js
assert(document.querySelectorAll('div')?.[1]?.classList?.contains('divider'));
```

Ihr `.divider`-Element sollte sich innerhalb Ihres `.label`-Elements befinden.

```js
assert(document.querySelector('.label')?.querySelector('.divider'));
```

Ihr `.divider`-Element sollte nach Ihrem `h1`-Element kommen.

```js
assert(document.querySelector('.divider')?.previousElementSibling?.localName === 'h1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
--fcc-editable-region--
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
--fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
}

p {
  margin: 0;
}
```
