---
id: 660ef857f2806aa626d29d17
title: 17단계
challengeType: 1
dashedName: step-17
---

# --description--

뺄셈(`-`), 곱셈(`*`), 나눗셈(`/`)도 수행할 수 있습니다. `console.log`에서 연산자와 숫자를 자유롭게 실험해 보세요. 다음 단계로 넘어갈 준비가 되면 `console.log`를 제거하세요.

# --hints--

코드에 `console.log()`가 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

`character` 변수는 여전히 `"Hello"` 값을 가져야 합니다.

```js
assert.equal(character, "Hello");
```

`count` 변수는 여전히 `8` 값을 가져야 합니다.

```js
assert.equal(count, 8);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let count = 8;
console.log(count + 1);
--fcc-editable-region--
```
