---
id: 660ef19b95d3308e7dd31bb6
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

<dfn>콘솔</dfn>은 JavaScript 출력 결과를 출력하고 확인할 수 있게 해줍니다. `console.log()`를 사용해 콘솔로 정보를 보낼 수 있습니다. 예를 들어, 이 코드는 `"Naomi"`을 콘솔에 출력합니다:

```js
let developer = "Naomi";
console.log(developer);
```

위 코드는 `developer` 안에 있는 이름으로 `console.log()` 변수를 접근합니다. 소괄호 안의 값이 출력될 값임을 기억하세요.

`character` 변수의 값을 콘솔에 출력하세요. 그런 다음 "Console" 버튼을 클릭해 JavaScript 콘솔을 확인하세요.

# --hints--

`character` 변수 초기화는 변경하지 마세요.

```js
assert.strictEqual(character, 'Hello');
```

코드에서 `console`에 접근해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

`log`의 `console` 속성에 접근해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

`.log()` 메서드를 호출하려면 소괄호를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

`character` 변수를 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
