---
id: 660ef55dd468079679ee0092
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

다음 변수를 선언할 준비가 되었습니다. `console.log` 문장 두 개와 `character` 재할당을 제거하세요.

또한 `secondCharacter` 변수를 제거하되, `character` 초기화는 변경하지 마세요.

# --hints--

코드에 `console.log()` 문장이 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

`character`의 값을 재할당하지 마세요.

```js
assert.isAtMost(__helpers.removeJSComments(code).match(/character\s*=/g).length, 1);
```

`character` 변수는 `"Hello"` 값을 가져야 합니다.

```js
assert.equal(character, "Hello");
```

`secondCharacter` 변수가 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /secondCharacter/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
secondCharacter = character;
console.log(secondCharacter);
--fcc-editable-region--
```
