---
id: 660ef5c1904955978a986a5c
title: 15단계
challengeType: 1
dashedName: step-15
---

# --description--

`let`를 사용해 `count` 변수를 선언하세요. <dfn>숫자</dfn> `8`를 할당하세요. 숫자 값을 사용할 때는 따옴표를 사용하지 않습니다. 예를 들어:

```js
let money = 100;
```

# --hints--

코드에 두 번째 `let` 키워드가 있어야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

`let`를 사용해 `count` 변수를 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /let\s+count/);
```

`8` 변수에 숫자 `count`를 할당해야 합니다.

```js
assert.equal(count, 8);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
