---
id: 660ef31a5be625914a0102cd
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

`let` 키워드로 변수를 선언하면 나중에 그 변수의 값을 <dfn>다시 할당</dfn>할 수 있습니다. 이 예제에서는 `programmer`의 값이 `"Naomi"`에서 `"CamperChan"`으로 변경됩니다.

```js
let programmer = "Naomi";
programmer = "CamperChan";
```

변수를 다시 할당할 때는 `let` 키워드를 다시 사용하지 않는다는 점을 기억하세요.

`console.log` 후에 `"World"` 변수에 `character` 값을 할당하세요.

# --hints--

코드에서 `character`를 두 번 사용해야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character/g), 3);
```

`character`를 다시 할당하려면 대입 연산자를 사용하세요.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\s*\=/g), 2);
```

코드에 문자열 `"World"`을 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /("|')World\1/);
```

`character` 변수는 `"World"` 값을 가져야 합니다.

```js
assert.equal(character, "World");
```

다시 할당할 때 `let`을 사용하지 마세요.

```js
assert.isBelow(__helpers.removeJSComments(code).match(/let/g).length, 2);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);

--fcc-editable-region--
```
