---
id: 660eebd83100d37862268781
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

현재 `character` 변수에는 값이 없습니다. <dfn>대입 연산자</dfn> `=`를 사용해 값을 할당할 수 있습니다. 예를 들어:

```js
let hello = "Hello";
```

변수를 선언할 때 값을 할당하는 것을 <dfn>초기화</dfn>라고 합니다.

선언하는 동안 `character` 변수에 값 `"Hello"`을 할당하여 초기화하세요.

# --hints--

대입 연산자 `=`를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character\s*=/);
```

문자열 `"Hello"`를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /('|")Hello\1/);
```

`"Hello"` 문자열에는 큰따옴표를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

`"Hello"`을 `character` 변수에 할당하세요.

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character;
--fcc-editable-region--
```
