---
id: 660ef0f7c4b8e68ccd1f0786
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

JavaScript에는 일곱 가지 원시 데이터 유형이 있으며, 그중 하나가 문자열입니다. JavaScript에서 <dfn>문자열</dfn>은 문자 수열을 나타내며, 작은따옴표(`'`)나 큰따옴표(`"`) 중 하나로 감쌀 수 있습니다.

문자열은 <dfn>불변</dfn>이라는 점에 유의하세요. 즉, 한 번 생성되면 변경할 수 없습니다. 변수에는 여전히 다른 값을 다시 할당할 수 있습니다.

`"Hello"` 문자열을 작은따옴표로 바꾸세요.

# --hints--

코드에 큰따옴표가 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /"/);
```

`"Hello"` 문자열에 작은따옴표를 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /'Hello'/);
```

코드에서 여전히 `let`를 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

여전히 `character` 변수를 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

`character` 변수에는 여전히 문자열 `"Hello"`이 값으로 있어야 합니다.

```js
assert.equal(character, "Hello");
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
--fcc-editable-region--
```
