---
id: 660ee6e3a242da6bd579de69
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

프로그래밍에서 가장 중요한 개념 중 하나는 변수입니다. <dfn>변수</dfn>는 값을 저장하는 특정 메모리 주소를 가리킵니다. 변수에는 이름이 부여되며, 이 이름을 사용해 코드 전체에서 해당 값에 접근할 수 있습니다.

변수를 선언한다는 것은 이름을 부여하는 것을 의미합니다. JavaScript에서는 보통 `let` 키워드를 사용해 선언합니다. 예를 들어, `hello` 변수를 선언하는 방법은 다음과 같습니다.

```js
let hello;
```

변수 이름은 특정 규칙을 따릅니다: 이름에는 문자, 숫자, 달러 기호, 밑줄을 포함할 수 있지만, 공백을 포함할 수 없고 숫자로 시작해서는 안 됩니다.

`let` 키워드를 사용해 `character`라는 변수를 선언하세요.

_참고_: JavaScript에서는 문장 끝에 세미콜론을 붙이는 것이 일반적인 관례입니다. `;`

# --hints--

코드에서 `let`를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

코드에서 `character`를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /character/);
```

`let`를 사용해 `character` 변수를 선언하세요.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

선언문은 세미콜론으로 끝나야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
