---
id: 660ef6355e8f5a9e67fe5f46
title: 16단계
challengeType: 1
dashedName: step-16
---

# --description--

`number` 데이터 타입으로 덧셈 같은 수학 연산을 할 수 있습니다. `count + 1`을 콘솔에 출력해 보세요.

# --hints--

코드에서 `console`에 접근해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

`.log()`의 `console`에 접근해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

`count + 1`를 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*count\s*\+\s*1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
