---
id: 660ef5105b8ba095307a0e50
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

이제 `character` 변수를 다시 콘솔에 출력하세요. 콘솔에서 `"Hello"` 문자열과 `"World"` 문자열이 차례로 보여야 합니다.

# --hints--

`console.log`를 두 번째로 사용해야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

`character`를 콘솔에 두 번째로 출력해야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*character\s*\)/g), 2);
```


새로운 `console.log()`는 재할당 뒤에 와야 합니다.

```js
const lines = __helpers.removeJSComments(code).split(/\n+/)
const reassign = lines.findIndex(l => l.match(/character\s+=\s+("|')World\1/));
const secondLog = lines.findLastIndex(l => l.match(/console\.log/));
assert.isBelow(reassign, secondLog);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";

--fcc-editable-region--
```
