---
id: 660f359af3e32e0f1a6880b7
title: 68단계
challengeType: 1
dashedName: step-68
---

# --description--

이제 콘솔에서 같은 문자들이 보일 거예요. `padRow` 함수는 이전에 하던 것과 똑같은 일을 하지만, 이제는 재사용할 수 있는 별도의 섹션에 들어 있습니다.

덧셈 연산자를 사용해 반복된 `" "` 문자열의 앞과 뒤에 공백 하나 `character`를 연결하세요.

문자열을 이렇게 연결할 때 `+` 연산자를 사용할 수 있다는 점을 기억하세요:

```js
" " + "string"
```

# --hints--

반환값의 앞에 공백 하나를 연결해야 합니다.

```js
assert.match(padRow(1, 1), /^\s/);
```

반환값의 뒤에 공백 하나를 연결해야 합니다.

```js
assert.match(padRow(1, 1), /\s$/);
```

`padRow()` 함수는 앞뒤에 공백이 있는 반복된 `character` 문자열을 반환해야 합니다.

```js
assert.equal(padRow(1, 1), " # ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
