---
id: 660f0da9bf1035c9097af20a
title: 28단계
challengeType: 1
dashedName: step-28
---

# --description--

콘솔에 `"freeCodeCamp"`가 출력된 것을 보셨을 거예요. 이는 `.pop()`이 배열에서 제거된 값을 반환하기 때문입니다. 그리고 이전에 `"freeCodeCamp"`를 배열 끝에 추가했죠.

그런데 `.push()`은 무엇을 반환할까요? 기존 `rows.push()`를 새로운 `pushed` 변수에 할당하고, 이를 출력하세요.

# --hints--

`pushed` 변수를 선언하세요.

```js
assert.match(__helpers.removeJSComments(code), /pushed/);
```

`let`를 사용해 `pushed` 변수를 선언하세요.

```js
assert.match(__helpers.removeJSComments(code), /let\s+pushed/);
```

`rows.push("freeCodeCamp")`를 `pushed` 변수에 할당하세요.

```js
assert.equal(pushed, 4);
```

`pushed` 변수를 출력하세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*pushed\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
rows.push("freeCodeCamp");

--fcc-editable-region--
let popped = rows.pop();
console.log(popped);
console.log(rows);
```
