---
id: 660f5179b3b0ca558f6b4d4f
title: 113단계
challengeType: 1
dashedName: step-113
---

# --description--

`if` 문을 사용해 `inverted`이 참인지 확인하세요. 여기서는 동등 연산자를 사용할 필요가 없다는 점을 기억하세요.

# --hints--

`if` 문을 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(/);
```

`if` 문은 `inverted`가 `true`인지 확인해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted/);
```

`if` 조건에는 어떤 비교 연산자도 사용하지 않아야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {

  rows.unshift(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
