---
id: 660f207334fabaeac3269c38
title: 41단계
challengeType: 1
dashedName: step-41
---

# --description--

`result` 문자열을 조작하려면 다른 유형의 반복문을 사용해야 합니다. 구체적으로, 반복 가능한 객체의 각 항목을 순차적으로 반복하며 임시로 변수에 할당하는 `for...of` 반복문입니다.

`for...of` 반복문의 구문은 다음과 같습니다:

```js
for (const value of iterable) {

}
```

변수는 전체 반복문 동안 존재하는 것이 아니라 단일 반복 단계 동안만 존재하므로 `const`을 사용할 수 있다는 점에 유의하세요.

`for...of` 반복문을 만들어 `rows` 배열을 순차적으로 반복하며 각 값을 `row` 변수에 할당하세요.

# --hints--

다른 `for` 키워드를 사용해야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for/g), 2);
```

`row` 변수를 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /\s+row\s+/);
```

`const`을 사용해 `row` 변수를 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /const\s+row\s+/);
```

`for...of` 반복문이 `row` 변수를 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+/);
```

`row` 변수는 `rows`에서 `of` 키워드를 사용해 추출해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)/);
```

`for...of` 반복문 본문은 비워 두어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--

console.log(result);
```
