---
id: 660f4a1472f8e63d76162ce5
title: 99단계
challengeType: 1
dashedName: step-99
---

# --description--

피라미드를 거꾸로, 즉 <dfn>뒤집힌</dfn> 모양으로 만든다면 어떨까요? 직접 시도해 보세요!

새로운 `for` 반복문을 만들어 시작하세요. 반복 변수 `i`를 선언하고 `count` 값을 할당한 다음, 조건과 반복문에는 불리언 `false`을 사용하세요.

# --hints--

코드에 `for` 반복문이 있어야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for\s*\(/g), 2);
```

`for` 반복문은 `i` 값을 사용해 `count`를 초기화해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count/);
```

`for` 반복문은 조건으로 `false`를 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*false/);
```


`for` 반복문은 반복문 실행에 `false`를 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*false\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
