---
id: 660f4cde8dd305450514a1cb
title: 105단계
challengeType: 1
dashedName: step-105
---

# --description--

다음 방법을 준비하려면 이 반복문도 여러 줄 주석으로 주석 처리하세요.

# --hints--

`for` 반복문은 여러 줄 주석으로 주석 처리되어야 합니다.

```js
assert.lengthOf(code.match(/\/\*\s*for\s*\(/g), 2);
assert.lengthOf(__helpers.removeJSComments(code).match(/for\s*\(/g), 1);
```

`for` 반복문 본문도 여러 줄 주석의 일부로 주석 처리되어야 합니다.

```js
assert.notMatch(code, /\/\/\s*rows\.push/);
assert.lengthOf(code.match(/rows\.push/g), 3)
assert.notMatch(__helpers.removeJSComments(code), /rows\.push/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
