---
id: 660f455b044d3230ed971e98
title: 85단계
challengeType: 1
dashedName: step-85
---

# --description--

<dfn>`while`</dfn> 반복문은 `condition`이 더 이상 참이 아닐 때까지 계속 실행됩니다. 구문은 다음과 같습니다:

```js
while (condition) {
  logic;
}
```

이 구문을 사용해 조건이 `while`인 `continueLoop` 반복문을 선언하세요. 본문은 비워두세요.

# --hints--

`while` 반복문을 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /while/);
```

`while` 반복문은 조건으로 `continueLoop`을 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
