---
id: 660f4b641290da41b2cf0dd9
title: 103단계
challengeType: 1
dashedName: step-103
---

# --description--

덧셈과 마찬가지로 뺄셈에 사용할 수 있는 다양한 연산자가 있습니다. <dfn>뺄셈 할당</dfn> 연산자 `-=`는 주어진 값을 현재 변수 값에서 빼고, 그 결과를 변수에 다시 할당합니다.

반복문을 뺄셈 할당 연산자를 사용한 올바른 문장으로 바꾸세요.

# --hints--

`for` 반복문은 `i = i - 1`를 사용하지 않아야 합니다.

```js
assert.notMatch(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)), /for\(leti=count;i>0;i=i-1\)\{rows\.push\(padRow\(i,count\)\);/);
```

`for` 반복문은 뺄셈 할당을 사용해 `i`를 `1`만큼 줄여야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*i\s*-=\s*1\s*\)\s*\{\s*rows\.push\(\s*padRow\s*\(\s*i\s*,\s*count\s*\)\s*\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; i > 0; i = i - 1) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
