---
id: 660f3ba3cceef11b6ba08b59
title: 75단계
challengeType: 1
dashedName: step-75
---

# --description--

안타깝게도 이제 피라미드의 아래쪽이 사라졌습니다. 이는 또 다른 <dfn>경계값 오류</dfn>를 만들었기 때문입니다.

원래 루프는 `i` 값이 `0`부터 `7`까지였는데, `count`가 `8`이고 조건이 `i`가 `count`보다 작아야 하기 때문입니다. 이제 루프는 `i` 값이 `1`부터 `7`까지 실행되고 있습니다.

루프가 `i`이 `8`일 때도 실행되도록 업데이트해야 합니다. 논리를 보면, 루프가 `i`가 `count`보다 <dfn>작거나 같다</dfn>일 때 실행되어야 한다는 뜻입니다. 이를 위해 작거나 같다 연산자 `<=`를 사용할 수 있습니다.

루프 조건을 `i`가 `count`보다 작거나 같을 때 실행되도록 업데이트하세요.

# --hints--

`for` 루프는 `i`가 `count`보다 작은지 확인하면 안 됩니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

`for` 루프는 `i`가 `count`보다 작거나 같은지 확인해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<=\s*count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
