---
id: 660f34e99571070d56d2f231
title: 67단계
challengeType: 1
dashedName: step-67
---

# --description--

`padRow` 함수에는 두 개의 매개변수가 정의되어 있습니다. 함수가 호출될 때 해당 매개변수에 값이 전달됩니다.

함수 호출에 제공하는 값을 <dfn>인수</dfn>라고 하며, 함수 호출에 인수를 <dfn>전달</dfn>한다고 합니다. 다음은 `"Hello"`이 인수로 전달된 함수 호출 예입니다:

```js
test("Hello");
```

`i + 1`와 `count`을 `padRow` 호출에 인수로 전달하세요. 매개변수처럼 인수도 쉼표로 구분합니다.

# --hints--

`i + 1`를 `padRow()` 호출에 전달해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1/);
```

`i + 1` 인수 뒤에 쉼표가 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*/);
```

두 번째 인수로 `count`를 전달해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*count\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow())
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
