---
id: 660f2fbd45b520046cac68e8
title: 65단계
challengeType: 1
dashedName: step-65
---

# --description--

이전 단계에서 반환값에 대해 배운 것을 기억하세요. 함수는 애플리케이션이 별도로 사용할 수 있도록 값을 <dfn>반환</dfn>할 수 있습니다.

함수에서 `return` 키워드는 반환값을 지정하는 데 사용됩니다. 예를 들어, 이 함수는 첫 번째 매개변수로 전달된 값을 반환합니다:

```js
function name(parameter) {
  return parameter;
}
```

`return` 키워드를 사용해 `character` 변수를 `rowNumber` 번 반복한 값을 반환하세요.

# --hints--

`.repeat()` 메서드를 사용하세요.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/\.repeat\(/g), 2);
```

`.repeat()` 변수에 `character` 메서드를 사용하세요.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\.repeat\(/g), 2);
```

`rowNumber` 호출에 `.repeat()` 을 전달하세요.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*rowNumber\s*\)/);
```

`return` 키워드를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /return/);
```

`.repeat()` 호출 결과를 반환하세요.

```js
assert.equal(padRow(3), "###");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
