---
id: 660f1cedf3676fe26122ebf6
title: 38단계
challengeType: 1
dashedName: step-38
---

# --description--

콘솔에 0부터 7까지 숫자가 한 줄에 하나씩 출력되는 것을 확인하세요. 이것이 피라미드를 생성하는 기반이 될 것입니다.

로그 문장을 `i` 배열에 `rows`을 추가하는 문장으로 바꾸세요.

# --hints--

`console.log` 호출이 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

`.push()` 배열에 `rows`을 호출해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

`i` 배열에 `rows`을 추가해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

`.push()`은 `for` 반복문 안에서 일어나야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*rows\.push\(\s*i\s*\);?\s*\}/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  console.log(i);
}
--fcc-editable-region--
```
