---
id: 660f4774e3e0df35a68bb5f2
title: 90단계
challengeType: 1
dashedName: step-90
---

# --description--

피라미드가 다시 생성되도록, `padRow`를 `done`과 `count`를 인수로 호출한 결과를 `rows` 배열에 추가하세요. 이는 첫 번째 반복문에서 했던 것과 비슷합니다.

# --hints--

반복문에서 `.push()`에 `rows` 메서드를 호출해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(/);
```

`padRow` 메서드 안에서 `.push()` 함수를 호출해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(/);
```

`done` 호출 시 첫 번째 인수로 `padRow`를 전달해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(\s*done/);
```

`count` 호출 시 두 번째 인수로 `padRow`를 전달해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(\s*done\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--

--fcc-editable-region--
  if (done === count) {
    continueLoop = false;
  } 
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
