---
id: 660f23b53db70af0f2620e78
title: 43단계
challengeType: 1
dashedName: step-43
---

# --description--

이제 모든 숫자가 한 줄에 나타납니다. 이렇게 하면 피라미드를 만들 수 없습니다.

각 행에 새 줄을 추가해야 합니다. 하지만 JavaScript에서 따옴표 사이에 줄 바꿈을 넣으려고 리턴 키를 누르면 구문 오류가 발생합니다. 대신 문자열이 출력될 때 새 줄로 해석되는 특별한 <dfn>이스케이프 시퀀스</dfn> `\n`를 사용해야 합니다. 예를 들어:

```js
lineOne = lineOne + "\n" + lineTwo;
```

두 번째 덧셈 연산자를 사용해 `result`와 `row` 값 뒤에 새 줄을 추가하세요.

# --hints--

`\n` 이스케이프 시퀀스를 사용하세요. 문자열이어야 하므로 따옴표로 감싸져야 한다는 점을 기억하세요.

```js
assert.match(__helpers.removeJSComments(code), /('|")\\n\1/);
```

`row` 변수를 `result` 변수에 연결하세요.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row\s*\+\s*('|")\\n\1/);
```

`\n` 이스케이프 시퀀스를 `row` 변수에 연결하세요.

```js
assert.match(__helpers.removeJSComments(code), /row\s*\+\s*('|")\\n\1/);
```

전체 연결 결과를 `result` 변수에 다시 할당하세요.

```js
assert.strictEqual(result, "0\n1\n2\n3\n4\n5\n6\n7\n");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
