---
id: 660f374d532dc41189cc9cc2
title: 69단계
challengeType: 1
dashedName: step-69
---

# --description--

이제 약간의 수학을 해볼 시간입니다. 3단 피라미드를 생각해 보세요. 가운데 정렬하려면 다음과 비슷하게 보일 것입니다:

```md
··#··
·###·
#####
```

빈 칸은 가독성을 위해 중간 점(인터펑트)으로 교체되었습니다. 이 패턴을 확장해 보면, 각 행의 시작과 끝에 있는 공백이 일정한 패턴을 따른다는 것을 알 수 있습니다.

빈 칸 문자열을 `rowCount - rowNumber`번 반복하도록 업데이트하세요.

결과를 확인하려면 콘솔을 열어 보세요.

# --hints--

`.repeat()` 문자열에 `" "`을 호출하여 `rowCount - rowNumber`번 반복해야 합니다.

```js
assert.equal(padRow(1, 3), "  #  ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " " + character.repeat(rowNumber) + " ";
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
