---
id: 660f53ad3d39175c5d4335ac
title: 117단계
challengeType: 1
dashedName: step-117
---

# --description--

피라미드 생성기가 원래 계획보다 더 많은 기능을 갖춘 완성된 상태입니다! 다음 단계는 코드를 정리하는 것입니다.

코드에서 모든 주석, 한 줄 주석과 여러 줄 주석을 모두 제거하세요.

# --hints--

코드에 한 줄 주석이 없어야 합니다.

```js
assert.notMatch(code, /\/\//);
```

코드에 여러 줄 주석이 없어야 합니다.

```js
assert.notMatch(code, /(?:\*\/|\/\*)/);
```

코드에 주석이 없어야 합니다.

```js
assert.equal(code, __helpers.removeJSComments(code));
```

여러 줄 주석으로 주석 처리된 코드를 제거해야 합니다.

```js
assert.notMatch(code, /while/);
assert.lengthOf(code.match(/for\s*\(/g), 2)
assert.lengthOf(code.match(/rows\.push/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
