---
id: 660f48a419b40238e2b8b4d5
title: 93단계
challengeType: 1
dashedName: step-93
---

# --description--

루프가 더 이상 `continueLoop` 변수를 사용하지 않습니다. 이로 인해 해당 변수가 <dfn>사용하지 않는 선언</dfn>이 되었습니다. 일반적으로, 미래의 혼란을 막기 위해 사용하지 않는 선언은 피하는 것이 좋습니다.

`continueLoop` 변수를 제거하세요.

# --hints--

더 이상 `continueLoop` 변수가 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /continueLoop/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
