---
id: 660f1a00ac619ddc1e259a66
title: 35단계
challengeType: 1
dashedName: step-35
---

# --description--

`for` 반복문의 <dfn>조건</dfn>은 반복할 횟수를 알려줍니다. `condition`이 false가 되면 반복문이 멈춥니다.

JavaScript에서 불리언 값은 `true` 또는 `false`일 수 있습니다. 이것들은 문자열이 아니며, 차이점은 나중에 더 배우게 됩니다.

지금은 <dfn>작다</dfn> 연산자(`<`)를 사용할 것입니다. 이 연산자는 왼쪽 값이 오른쪽 값보다 작은지 확인할 수 있게 해줍니다. 예를 들어, `count < 3`는 `true`이 `count`일 때 `2`로 평가되고, `false`이 `count`일 때 `4`로 평가됩니다.

`"condition"` 문자열을 `i`이 `count`보다 작은지 확인하는 조건으로 바꾸세요.

# --hints--

작다 연산자를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /</);
```

`i`가 `count`보다 작은지 확인하려면 작다 연산자를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

`for` 반복문은 `i < count`을 조건으로 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*count\s*;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; "condition"; "iteration") {

}
--fcc-editable-region--
```
