---
id: 660f280dda5040f707c76b4a
title: 45단계
challengeType: 1
dashedName: step-45
---

# --description--

이제 `#` 문자들이 연속으로 있지만, 피라미드 모양은 아직 없습니다. 다행히도 `i` 변수는 반복문의 현재 "행" 번호를 나타내므로, 이를 사용해 피라미드 모양을 만들 수 있습니다.

이를 위해 문자열에 사용할 수 있는 `.repeat()` 메서드를 활용할 것입니다. 이 메서드는 인수로 숫자를 받아, 대상 문자열을 그 횟수만큼 반복합니다. 예를 들어, `.repeat()`를 사용해 `"Code! Code! Code!"` 문자열을 생성할 수 있습니다:

```js
const activity = "Code! ";
activity.repeat(3);
```

`.repeat()` 메서드를 `character`에 사용하고, 숫자로 `i`을 전달하세요.

# --hints--

`.repeat()` 메서드를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /\.repeat\(/);
```

`.repeat()` 메서드를 `character` 변수에 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(/);
```

`i`을 `.repeat()` 메서드에 인수로 전달하세요.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\)/)
```

`.repeat()` 메서드를 `.push()` 메서드 안에서 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\.repeat\(\s*i\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
