---
id: 660f4377a359972c521d3f4b
title: 79단계
challengeType: 1
dashedName: step-79
---

# --description--

`true`가 실제로 true이기 때문에 문자열이 콘솔에 출력되는 것을 볼 수 있습니다.

`if` 문장의 조건을 불리언 `false`로 변경하세요.

# --hints--

`if` 조건은 `false`를 조건으로 가져야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if (true) {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
