---
id: 660f447efc0e722f016c1be0
title: 81단계
challengeType: 1
dashedName: step-81
---

# --description--

텍스트가 다시 나타났습니다! 이는 `"false"`가 문자열이고, 불리언으로 평가할 때 `true`가 되기 때문입니다. 즉, `"false"`는 참 같은 값입니다.

<dfn>참 같은 값</dfn>은 불리언으로 평가할 때 참으로 간주되는 값입니다. JavaScript에서 만나는 대부분의 값은 참 같은 값입니다.

<dfn>거짓 같은 값</dfn>은 그 반대로, 불리언으로 평가할 때 거짓으로 간주되는 값입니다. JavaScript에는 정의된 거짓 같은 값 목록이 있습니다. 그중 일부는 `false`, `0`, `""`, `null`, `undefined`, 그리고 `NaN`입니다.

`if` 조건을 빈 문자열 `""`로 바꿔 보세요. 빈 문자열은 거짓 같은 값입니다.

# --hints--

`if` 문장은 조건으로 `""`를 가져야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*("|')\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("false") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
