---
id: 660f46b9c417a8341729a3ab
title: 89단계
challengeType: 1
dashedName: step-89
---

# --description--

`done`가 `count` 값에 도달하면 반복문 실행을 멈추게 하세요.

`if` 본문 안에서 불리언 `false`를 `continueLoop` 변수에 할당하세요.

# --hints--

`if` 본문에서 `false`를 `continueLoop` 변수에 할당해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{\s*continueLoop\s*=\s*false;?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done === count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
