---
id: 660f18f059fe0fda192ce394
title: 34단계
challengeType: 1
dashedName: step-34
---

# --description--

루프에 이제 적절한 반복 변수가 필요합니다. <dfn>반복자</dfn>는 루프가 어떻게 반복되거나 로직을 진행할지 제어하기 위해 `for` 루프 안에서 특별히 선언할 수 있는 변수입니다.

루프에서 반복자 변수로 `i`을 사용하는 것이 일반적인 관례입니다. `for` 루프는 소괄호 `()` 안에서 이를 선언할 수 있게 해줍니다. 예를 들어, 여기 `for` 루프가 `index` 변수를 선언하고 값 `100`을 할당하는 예가 있습니다.

```js
for (let index = 100; "second"; "third") {

}
```

문자열 `"iterator"`을 `let` 변수에 대한 `i` 선언으로 바꾸세요. 시작 값으로 `0`을 할당하세요. 이렇게 하면 루프가 **처음 실행될 때** `i` 변수에 값 `0`이 할당됩니다.

# --hints--

`let`을 사용해 `i` 변수를 선언하세요.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i/);
```

`0`을 `i` 변수에 할당하세요.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i\s*=\s*0/);
```

`for` 루프는 `i`에서 `0` 반복자를 시작해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for ("iterator"; "condition"; "iteration") {

}
--fcc-editable-region--
```
