---
id: 660f3915b41a441537ec9f5e
title: 72단계
challengeType: 1
dashedName: step-72
---

# --description--

`i`에 `1`만큼만 증가시키기 때문에 <dfn>증가 연산자</dfn> `++`를 사용할 수 있습니다. 이 연산자는 변수의 값을 1만큼 증가시켜 해당 변수에 대한 할당을 업데이트합니다. 예를 들어, 여기서 `test`는 `8`가 됩니다:

```js
let test = 7;
test++;
```

반복문의 반복 횟수에 대해 덧셈 할당 대신 증가 연산자를 사용하세요.

# --hints--

`for` 반복문에서는 `i`에 덧셈 할당을 사용하지 마세요.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

`for` 반복문에서는 `i`에 증가 연산자를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i += 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
