---
id: 660f4455f457ef2e3ec6920f
title: 80단계
challengeType: 1
dashedName: step-80
---

# --description--

이제 문자열이 더 이상 출력되지 않습니다. 왜냐하면 `false`가 `true`가 아니기 때문입니다. 그렇다면 다른 값들은 어떨까요?

조건을 문자열 `"false"`로 바꿔보세요.

# --hints--

`if` 문장은 조건으로 문자열 `"false"`를 가져야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*('|")false\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if (false) {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
