---
id: 660f38c34a4de6141c0c369f
title: 71단계
challengeType: 1
dashedName: step-71
---

# --description--

피라미드 생성기가 이제 예상대로 작동합니다. 하지만 작성한 코드를 더 깊이 탐구할 수 있는 좋은 기회입니다.

덧셈 연산자만이 변수에 값을 더하는 유일한 방법은 아닙니다. <dfn>덧셈 할당</dfn> 연산자는 "변수의 원래 값에 이 값을 더하고, 그 결과를 다시 변수에 할당한다"는 의미의 축약형으로 사용할 수 있습니다. 예를 들어, 다음 두 문장은 같은 결과를 냅니다:

```js
test = test + 1;
test += 1;
```

`for` 반복문 안의 반복문 문장을 덧셈 할당을 사용하도록 수정하세요.

# --hints--

`for` 반복문에서는 `i = i + 1`를 사용하지 않아야 합니다;

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

`for` 반복문에서는 `i`와 함께 덧셈 할당을 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
