---
id: 660f4efcb8068e4cb470dca1
title: 109단계
challengeType: 1
dashedName: step-109
---

# --description--

이 메서드들을 시도해봤으니, 이제 또 다른 역피라미드 방식을 시도할 수 있습니다. 하지만 먼저 실험 내용을 정리해야 합니다.

`numbers` 배열과 메서드 호출 및 로그 호출을 제거하세요.

# --hints--

`numbers` 배열이 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /numbers/);
```

`unshifted` 변수가 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /unshifted/);
```

`shifted` 변수가 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /shifted/);
```

`console.log` 문장이 없어야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];
const shifted = numbers.shift();
console.log(shifted);
const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
