---
id: 660f0c34aad72dc712b97624
title: 27단계
challengeType: 1
dashedName: step-27
---

# --description--

이 프로젝트에 꼭 필요한 또 다른 메서드는 `.pop()` 메서드입니다. 이 메서드는 배열에서 마지막 요소를 제거하고 그 요소를 <dfn>반환</dfn>합니다.

메서드가 값을 반환할 때, 그 값을 다시 돌려주는 것으로 생각할 수 있으며, 코드의 다른 부분에서 사용할 수 있게 됩니다.

`popped`이라는 새 변수를 만들고 `rows.pop()`의 결과를 할당하세요. 그런 다음 `popped`을 콘솔에 출력하세요.

# --hints--

`popped`라는 변수를 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /popped/);
```

`let`를 사용해 `popped`라는 변수를 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /let\s+popped/);
```

`.pop()` 메서드를 호출해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /\.pop\(\s*\)/);
```

`.pop()` 배열에서 `rows` 메서드를 호출해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /rows\.pop\(\s*\)/);
```

`popped` 변수를 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*popped\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];
rows.push("freeCodeCamp");

console.log(rows);
--fcc-editable-region--
```
