---
id: 660f4ae5b3924c3fc3373973
title: 101단계
challengeType: 1
dashedName: step-101
---

# --description--

반복문도 달라질 거예요. 각 반복마다 `1`에 `i`를 더하는 대신 `1`를 빼야 합니다.

이전에 `i = i + 1`에서 했던 것처럼, 반복문을 수정해서 `i`가 자기 자신에서 `1`를 뺀 값을 갖도록 하세요.

# --hints--

`for` 반복문은 반복 변수로 `i = i - 1`를 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*i\s*=\s*i\s*-\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; i > 0; false) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
