---
id: 660f47afe4c98536715d5fa4
title: 91단계
challengeType: 1
dashedName: step-91
---

# --description--

<dfn>엄격 부등호</dfn> 연산자 `!==`는 두 값이 같지 않거나 같은 유형이 아닌지 확인할 수 있게 해줍니다. 구문은 동등 연산자와 비슷합니다: `value !== 4`.

현재 `while` 반복문은 `continueLoop`이 true일 때만 실행됩니다. `while` 반복문 조건을 `done`가 `count`와 같지 않은지 확인하도록 업데이트하세요.

# --hints--

`while` 반복문은 `done`과 `count`가 같지 않은지 확인해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*(?:done\s*!==\s*count|count\s*!==\s*done)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;
  rows.push(padRow(done, count));
  if (done === count) {
    continueLoop = false;
  } 
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
