---
id: 660f3ce51f70571e1c5227c8
title: 76단계
challengeType: 1
dashedName: step-76
---

# --description--

주석은 코드가 특정 방식을 사용하는 이유를 설명하거나 미래의 자신을 위해 해야 할 일을 적어두는 데 도움이 됩니다.

JavaScript에서는 `//`를 사용해 코드에 한 줄 주석을 남길 수 있습니다.

코드를 다른 종류의 반복문으로 바꾸라는 메모를 남기기 위해 `for` 반복문 위에 한 줄 주석을 추가하세요.

# --hints--

한 줄 주석은 `//`로 시작해야 합니다.

```js
assert.match(code, /\/\//);
```

한 줄 주석은 최소 다섯 글자 이상이어야 합니다.

```js
assert.match(code, /\/\/.{5,}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--

for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
