---
id: 660f0ee51d7460ce88cd248d
title: 29단계
challengeType: 1
dashedName: step-29
---

# --description--

콘솔에서 `4`를 볼 것으로 예상했나요? `.push()`는 값을 추가한 후 배열의 새 길이를 반환합니다.

메서드가 어떤 값을 반환하는지 아는 것이 중요합니다. `.push()`와 `.pop()`를 사용해 실험해 보세요. 준비가 되면 모든 `.push()`와 `.pop()` 호출, 그리고 `console.log` 문장을 제거하세요.

# --hints--

`.push()` 호출이 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.push\(/);
```

`.pop()` 호출이 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.pop\(/);
```

로그 문장이 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log/);
```

`popped` 변수가 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /popped/);
```

`pushed` 변수가 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /pushed/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let pushed = rows.push("freeCodeCamp");
console.log(pushed);
let popped = rows.pop();
console.log(popped);
console.log(rows);
--fcc-editable-region--
```
