---
id: 660f07d231941bc11719f664
title: 21단계
challengeType: 1
dashedName: step-21
---

# --description--

배열은 <dfn>변경 가능</dfn>하다고 간주된다는 점에서 특별합니다. 즉, 인덱스의 값을 직접 변경할 수 있습니다.

예를 들어, 이 코드는 배열의 두 번째 요소에 숫자 `25`를 할당합니다:

```js
let array = [1, 2, 3];
array[1] = 25;
console.log(array); // prints [1, 25, 3]
```

`rows` 배열의 **세 번째** 요소를 숫자 `10`로 변경하세요. 그런 다음 `rows` 배열을 콘솔에 출력하세요.

# --hints--

`rows` 배열에 다시 대괄호 표기법을 사용하세요.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows\[/g), 2)
```

`rows` 배열의 세 번째 요소에 접근하세요.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]/);
```

`rows` 배열의 세 번째 요소에 대입 연산자를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]\s*=/);
```

`10` 배열의 세 번째 요소에 값 `rows`를 할당하세요.

```js
assert.equal(rows[2], 10);
```

코드에 두 번째 `console.log` 문장이 있어야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

`rows` 배열을 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--

--fcc-editable-region--
```
