---
id: 660f49e32001983c90b75850
title: 98단계
challengeType: 1
dashedName: step-98
---

# --description--

아주 깔끔하고 기능적인 반복문입니다. 잘했어요! 하지만 아직 더 탐구할 부분이 있습니다.

여러 줄 주석을 사용해 `while` 반복문을 주석 처리하세요.

# --hints--

`while` 반복문을 주석 처리해야 합니다.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

`while` 반복문 본문을 주석 처리해야 합니다.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /rows\.push/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
