---
id: 660f540c2176ea5dec01306d
title: 118단계
challengeType: 1
dashedName: step-118
---

# --description--

잘했어요! `character`, `count`, `inverted` 변수에 다양한 값을 실험해 보세요.

다음 프로젝트로 넘어갈 준비가 되면 `character`을 `"!"`로, `count`를 `10`로, `inverted`을 `false`로 설정해 계속 진행하세요.

첫 JavaScript 프로젝트를 완료한 것을 축하합니다!

# --hints--

`character`을 `"!"`으로 설정하세요.

```js
assert.equal(character, "!");
```

`count`을 `10`로 설정하세요.

```js
assert.equal(count, 10);
```

`inverted`을 `false`로 설정하세요.

```js
assert.equal(inverted, false);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "#";
const count = 8;
const rows = [];
let inverted = true;
--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

# --solutions--

```js
const character = "!";
const count = 10;
const rows = [];
let inverted = false;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
