---
id: 660f033cf051ebb50ea3bf48
title: 18단계
challengeType: 1
dashedName: step-18
---

# --description--

프로그래밍에서는 많은 데이터를 다뤄야 할 때가 많습니다. 데이터를 체계적으로 관리하고 정리하는 데 도움이 되는 다양한 자료 구조가 있습니다. 그중 가장 기본적인 자료 구조 중 하나가 배열입니다.

배열은 일련의 값을 담을 수 있는 비원시 데이터 유형입니다. 비원시 데이터 유형은 원시 데이터 유형과 달리 더 복잡한 데이터를 담을 수 있습니다. 문자열이나 숫자 같은 원시 데이터 유형은 한 번에 하나의 값만 담을 수 있습니다.

배열은 대괄호(`[]`)를 사용해 표시합니다. 다음은 빈 배열 값을 가진 변수의 예입니다.

```js
let array = [];
```

`rows` 변수를 선언하고 빈 배열을 할당하세요.

# --hints--

`rows` 변수가 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /rows/);
```

`let`를 사용해 `rows` 변수를 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /let\s+rows/);
```

`rows` 변수에 빈 배열을 할당해야 합니다.

```js
assert.deepEqual(rows, []);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
