---
id: 660f39b444fd6f16d1e49c1f
title: 73단계
challengeType: 1
dashedName: step-73
---

# --description--

`i + 1` 호출에 `padRow`를 전달하는 대신, 반복문을 `1`에서 시작할 수 있습니다. 이렇게 하면 1부터 시작하는 인덱스 반복문을 만들 수 있습니다.

반복자를 `1` 대신 `0`에서 시작하도록 업데이트하세요.

# --hints--

`for` 반복문은 `i`에서 `1`를 초기화해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
