---
id: 660f0a55847d6cc485f29ba5
title: 26단계
challengeType: 1
dashedName: step-26
---

# --description--

JavaScript에서 <dfn>메서드</dfn>는 특정 값이나 객체와 연관된 함수입니다. 이미 접한 예로는 `.log()` 객체의 일부인 `console` 메서드가 있습니다.

배열에는 자체 메서드가 있으며, 처음 탐색할 메서드는 `.push()` 메서드입니다. 이 메서드는 배열 끝에 값을 "추가(push)"할 수 있게 해줍니다. 다음은 숫자 `12`을 배열 끝에 추가하는 예입니다:

```js
array.push(12);
```

`.push()`를 사용해 문자열 `"freeCodeCamp"`를 `rows` 배열 끝에 추가하세요. 배열에 변경 사항이 반영된 것을 볼 수 있도록 `console.log` 이전에 이 코드를 추가하세요.

# --hints--

코드에서 `.push()`를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /\.push\(/);
```

`.push()` 배열의 `rows` 메서드를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

문자열 `"freeCodeCamp"`을 `.push()` 메서드에 전달하세요.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*('|")freeCodeCamp\1\s*\);?/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];

console.log(rows);
--fcc-editable-region--
```
