---
id: 660f415b76859a2736771607
title: 78단계
challengeType: 1
dashedName: step-78
---

# --description--

피라미드가 다시 사라졌습니다. 괜찮아요 — 예상된 일입니다.

새로운 반복문을 만들기 전에 `if` 문에 대해 배워야 합니다. <dfn>`if` 문</dfn>은 조건이 충족될 때만 코드 블록을 실행할 수 있게 해줍니다. 다음 구문을 사용합니다:

```js
if (condition) {
  logic
}
```

불리언 `if`를 조건으로 하는 `true` 문을 만드세요. 본문에서는 문자열 `"Condition is true"`를 출력하세요.

# --hints--

`if` 문을 만들어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /if/);
```

`if` 문은 `true`를 조건으로 해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)/);
```

`if` 본문은 `"Condition is true"`를 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)\s*\{\s*console\.log\(\s*("|')Condition is true\1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
