---
id: 660f4a83373de83ea101685f
title: 100단계
challengeType: 1
dashedName: step-100
---

# --description--

반대 방향으로 반복할 것이므로, 반복문은 `i`가 `0`보다 클 때 실행되어야 합니다. 이를 위해 <dfn>보다 큰</dfn> 연산자 `>`를 사용할 수 있습니다.

반복문의 조건을 `i`가 `0`보다 클 때 실행되도록 설정하세요.

# --hints--

`for` 반복문은 `i`가 `0`보다 클 때 실행되어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; false; false) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
