---
id: 660f20473aef47e9b8c9afc6
title: 40단계
challengeType: 1
dashedName: step-40
---

# --description--

`result`의 값을 출력하는 로그 문을 추가하세요. 사용하는 콘솔에 따라 아무것도 출력되지 않을 수도 있습니다.

# --hints--

`console.log` 호출을 추가해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

`result` 변수를 로그로 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--
```
