---
id: 660f4934fb48f63abd5ae371
title: 95단계
challengeType: 1
dashedName: step-95
---

# --description--

`done`를 사용해 생성된 행 수를 추적하는 것은 기능적으로 가능하지만, 실제로는 로직을 좀 더 깔끔하게 정리할 수 있습니다.

배열에는 배열에 몇 개의 값 또는 <dfn>요소</dfn>가 있는지 확인할 수 있는 특별한 `length` 속성이 있습니다. 이 속성은 `myArray.length`와 같은 구문으로 접근할 수 있습니다.

`rows.length` 호출에서 `padRow`을 사용하면 경계값 오류가 발생하는데, 그 이유는 `done`가 호출 *전에* 증가하기 때문입니다.

조건을 `rows.length`가 `count`보다 작은지 확인하도록 업데이트하세요.

# --hints--

`while` 반복문은 `rows.length`가 `count`보다 작은지 확인해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*rows\.length\s*<\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done <= count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
