---
id: 660f1b6e60bd9edf902c81fd
title: 36단계
challengeType: 1
dashedName: step-36
---

# --description--

<dfn>반복</dfn>문은 각 실행 후 반복자에 대해 루프가 무엇을 할지 알려줍니다.

변수를 다시 할당할 때, 다시 할당하기 전의 이전 값을 참조하는 데 변수를 사용할 수 있습니다. 이를 통해 기존 숫자에 3을 더하는 것 같은 작업을 할 수 있습니다. 예를 들어, `bees = bees + 3`는 `bees`의 값을 3만큼 증가시킵니다.

그 구문을 사용해 `"iteration"` 문자열을 `i`을 1만큼 증가시키는 다시 할당문으로 바꾸세요.

# --hints--

`i` 변수에 1을 더해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\s*1/);
```

`i + 1`를 다시 `i` 변수에 할당해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

`for` 루프는 각 반복마다 `i`을 `1`만큼 증가시켜야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; "iteration") {

}
--fcc-editable-region--
```
