---
id: 660f530d6e33d159e1bf4947
title: 115단계
challengeType: 1
dashedName: step-115
---

# --description--

피라미드가 거꾸로 되어 있지 않다면, 일반적인 순서로 피라미드를 만드는 `else` 블록이 필요합니다.

이전 단계에서, 다음과 같은 `else` 문을 다루는 방법을 배웠습니다:

```js
if (condition) {
  // if condition is true, run this code
} else {
  // if condition is false, run this code
}
```

`else` 블록에 `if` 블록을 추가하세요.

# --hints--

`else` 블록을 추가해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted\s*\)\s*\{\s*rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)\s*\);\s*\}\s*else\s*\{/);
```

`else` 블록은 비어 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted\s*\)\s*\{\s*rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)\s*\);\s*\}\s*else\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  }

}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
