---
id: 660f44f10ea40f300b896a5e
title: 84단계
challengeType: 1
dashedName: step-84
---

# --description--

`if...else if...else` 문을 다뤄보았으니 이제 코드에서 제거하세요.

그다음 `let`을 사용해 `continueLoop` 변수를 선언하고 불리언 `false` 값을 할당하세요. 그리고 `let`를 사용해 `done` 변수를 선언하고 값 `0`을 할당하세요.

# --hints--

`if` 문이 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /if\s*\(/);
```

`else if` 문이 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

`else` 문이 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s*\{/);
```

`let`을 사용해 `continueLoop` 변수를 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /let\s+continueLoop/);
```

`continueLoop` 변수는 `false` 값을 가져야 합니다.

```js
assert.isFalse(continueLoop);
```

`let`를 사용해 `done` 변수를 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /let\s+done/);
```

`done` 변수는 `0` 값을 가져야 합니다.

```js
assert.strictEqual(done, 0);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
} else {
  console.log("This is the else block");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
