---
id: 660f4de78f775e480ba2e451
title: 107단계
challengeType: 1
dashedName: step-107
---

# --description--

배열의 `.unshift()` 메서드는 `.push()`과 달리 배열의 끝이 아닌 **처음**에 값을 추가할 수 있습니다. `.unshift()`는 호출된 배열의 새로운 길이를 반환합니다.

```js
const countDown = [2, 1, 0];
const newLength = countDown.unshift(3);
console.log(countDown); // [3, 2, 1, 0]
console.log(newLength); // 4
```

`const`를 사용해 `unshifted` 변수를 선언하고, `.unshift()` 배열에서 `numbers`를 호출한 결과를 할당하세요. 인수로 `5`를 전달하세요. 그런 다음 `unshifted` 변수를 출력하세요.

# --hints--

`unshifted` 변수를 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /const\s+unshifted/);
```

`.unshift()` 배열에서 `numbers` 메서드를 호출해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /numbers\.unshift\(/);
```

`5` 호출에 `.unshift()`를 인수로 전달해야 합니다.

```js
assert.deepEqual(numbers, [5,1,2,3]);
```

`.unshift()` 호출 결과를 `unshifted` 변수에 할당해야 합니다.

```js
assert.equal(unshifted, 4);
```

`unshifted` 변수를 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*unshifted\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];

console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
