---
id: 660f4990b1caa03b9dc97a43
title: 97단계
challengeType: 1
dashedName: step-97
---

# --description--

이제 `done` 변수가 더 이상 필요하지 않습니다. 반복문에서 증가 연산을 제거하고 `done` 변수 선언도 삭제하세요.

# --hints--

`done` 변수를 증가시키면 안 됩니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /done\+\+/);
```

더 이상 `done` 변수가 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /done/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (rows.length < count) {
  done++;
  rows.push(padRow(rows.length + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
