---
id: 660f17294346b7d69e79db3d
title: 단계 32
challengeType: 1
dashedName: step-32
---

# --description--

이제 피라미드 생성기를 만들 준비가 되었습니다. `character` 변수는 피라미드의 기본 블록 역할을 할 것입니다.

`"Hello"`은(는) 그 용도에 적합하지 않을 수 있습니다. `character`의 값을 해시 기호 `"#"`로 변경하세요.

# --hints--

`character` 변수는 해시 기호여야 합니다.

```js
assert.equal(character, "#");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "Hello";
const count = 8;
const rows = [];
--fcc-editable-region--
```
