---
id: 660f3b664421471aa595170f
title: 74단계
challengeType: 1
dashedName: step-74
---

# --description--

피라미드가 지금 좀 이상해 보입니다. `1` 대신 `0`에서 루프를 시작하기 때문에 `i`에 전달할 때 `padRow`에 1을 더할 필요가 없습니다.

`padRow` 호출의 첫 번째 인수를 `i`로 업데이트하세요.

# --hints--

`i + 1`를 `padRow`에 전달하면 안 됩니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /padRow\(\s*i\s*\+\s*1/);
```

`i`를 `padRow`에 전달해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*i\s*,/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
