---
id: 660f50a21fe7645252804f2b
title: 112단계
challengeType: 1
dashedName: step-112
---

# --description--

뒤집힌 피라미드와 일반 피라미드 사이를 전환할 수 있는 방법이 있다면 어떨까요?

`inverted` 변수를 선언하고, 값 `true`을 할당하는 것부터 시작하세요. 이 변수는 코드에서 변경하지 않지만, 나중에 테스트가 수정할 수 있도록 `let`를 사용해야 합니다.

# --hints--

`inverted` 변수를 `let`로 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted/);
```

`inverted`를 값 `true`으로 초기화해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted\s*=\s*true;?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
--fcc-editable-region--

--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  rows.unshift(padRow(i, count));
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
