---
id: 660f2eccfe3f820304af1b39
title: 64단계
challengeType: 1
dashedName: step-64
---

# --description--

행을 어떻게 포맷할지 알기 위해서는, `padRow` 함수가 현재 몇 번째 행인지와 총 몇 행이 생성되는지 알아야 합니다.

가장 좋은 방법은 이들을 위한 함수 매개변수를 만드는 것입니다. `padRow` 함수에 `rowNumber`와 `rowCount` 매개변수를 추가하세요. 여러 매개변수는 쉼표로 구분합니다:

```js
function name(first, second) {

}
```

# --hints--

`padRow` 함수에 `rowNumber` 매개변수가 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber/);
```

`rowNumber` 매개변수 뒤에 쉼표를 추가하세요.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*/);
```

`padRow` 함수에 `rowCount` 매개변수가 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*rowCount\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
