---
id: 660f45ccf4ca5c31f253005a
title: 86단계
challengeType: 1
dashedName: step-86
---

# --description--

지금 `continueLoop`을 true로 바꾸면 `while` 루프가 무한히 실행됩니다. 이를 <dfn>무한 루프</dfn>라고 하며, 이런 상황을 피하도록 주의해야 합니다. 무한 루프는 시스템을 멈추게 하여 완전한 재시작이 필요할 수 있습니다.

이를 피하려면 루프 안에서 증가 연산자를 사용해 `done` 변수의 값을 증가시키는 것부터 시작하세요.

# --hints--

`while` 루프는 `done` 변수를 증가시켜야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
