---
id: 660f4cffb1459d45e34902d1
title: 106단계
challengeType: 1
dashedName: step-106
---

# --description--

사실, 이전처럼 "역방향"으로 반복하지 않고도 역피라미드를 만들 수 있습니다.

이를 위해서는 새로운 배열 메서드 몇 가지를 배워야 합니다. `const`를 사용해 `numbers` 변수를 선언하세요. `1`, `2`, `3` 요소를 가진 배열을 할당하세요. 그런 다음 `numbers` 배열을 출력하세요.

# --hints--

`const`를 사용해 `numbers` 변수를 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /const\s+numbers/);
```

`numbers` 변수는 배열이어야 합니다.

```js
assert.isArray(numbers);
```

`numbers` 배열은 `1`, `2`, `3` 요소를 그 순서대로 가져야 합니다.

```js
assert.deepEqual(numbers, [1,2,3]);
```

`numbers` 배열을 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*numbers\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
