---
id: 660f09a2694b59c3a10ee304
title: 23단계
challengeType: 1
dashedName: step-23
---

# --description--

우선, 첫 번째 콘솔 로그와 `rows[rows.length - 1]` 할당을 제거하세요. 두 번째 `rows` 로그 문장은 나중에 남겨두세요.

# --hints--

`console.log(rows[0])` 문장은 제거해야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*\]\s*\)/);
```

`rows[rows.length - 1]` 재할당은 제거해야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/);
```

`console.log(rows)` 문장은 제거하지 마세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
console.log(rows[0]);
rows[rows.length - 1] = 10;
--fcc-editable-region--
console.log(rows);
```
