---
id: 660f48e1d3682f39e81843c4
title: 94단계
challengeType: 1
dashedName: step-94
---

# --description--

피라미드 생성기는 여전히 작동하고 있습니다. 하지만 무한 루프에 빠질 가능성이 여전히 있습니다.

`done`가 `count`와 같지 않은지만 확인하고 있기 때문에, 만약 `done`가 `count`보다 **크다면** 루프가 영원히 계속될 수 있습니다.

루프 조건을 `done`가 `count`보다 작거나 같은지 확인하도록 업데이트하세요.

# --hints--

`while` 루프는 `done`가 `count`보다 작거나 같은지 확인해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*done\s*<=\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
