---
id: 660f0f980e98e8cf77f1ce31
title: 30단계
challengeType: 1
dashedName: step-30
---

# --description--

`rows` 선언을 다시 빈 배열로 할당하세요.

또한 `'Hello'` 문자열을 다시 큰따옴표로 사용하세요. 일반적으로 둘 중 어느 것을 선호하든 상관없지만, 프로젝트 전체에서 일관되게 사용하는 것이 좋습니다.

# --hints--

`rows` 배열은 비어 있어야 합니다.

```js
assert.empty(rows);
```

`"Hello"` 문자열은 큰따옴표를 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
```
