---
id: 660f17d4e9f227d86e834abd
title: 33단계
challengeType: 1
dashedName: step-33
---

# --description--

피라미드를 만들려면 여러 행을 생성해야 합니다. 조건이 충족될 때까지 작업을 반복해야 할 때는 <dfn>루프</dfn>를 사용합니다. 루프를 작성하는 방법은 여러 가지가 있습니다.

기본 `for` 루프부터 시작할 것입니다. `for` 루프는 다음 구문을 사용합니다:

```js
for (iterator; condition; iteration) {
  logic;
}
```

앞으로의 단계에서 루프의 각 컴포넌트를 자세히 살펴볼 것입니다. 지금은 `for` 루프를 만들어 세 가지 컴포넌트로 `"iterator"`, `"condition"`, `"iteration"`를 포함하세요. 루프 <dfn>본문</dfn>, 즉 중괄호 `{}` 안의 섹션은 비워 두세요.


# --hints--

`for` 루프가 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /for/);
```

`for` 루프의 첫 번째 컴포넌트는 문자열 `"iterator"`이어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1/);
```

`for` 루프의 두 번째 컴포넌트는 문자열 `"condition"`이어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2/);
```

`for` 루프의 세 번째 컴포넌트는 문자열 `"iteration"`이어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)/);
```

`for` 루프의 본문은 비워 두세요.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--
```
