---
id: 660f229d2dbe09ef2954a4a1
title: 42단계
challengeType: 1
dashedName: step-42
---

# --description--

이전 반복문에서 `+` 더하기 연산자를 사용해 `i` 값에 `1` 값을 더한 것을 기억하세요.

문자열 값에도 비슷한 작업을 할 수 있는데, 기존 문자열에 새 문자열을 덧붙이는 방식입니다. 예를 들어, `hello = hello + " World";` 는 `" World"` 변수에 저장된 기존 문자열에 `hello` 문자열을 추가합니다. 이것을 <dfn>연결</dfn>이라고 합니다.

`for...of` 반복문에서 더하기 연산자를 사용해 `row` 값에 `result` 값을 연결하세요.

# --hints--

`result` 변수에 연결 연산자를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /(?:result\s*\+|\+\s*result)/);
```

`row` 변수에 `result` 을 연결해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row|row\s*\+\s*result/);
```

연결 결과를 `result` 변수에 다시 할당하세요.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*(result\s*\+\s*row|row\s*\+\s*result);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {

}
--fcc-editable-region--

console.log(result);
```
