---
id: 660f1bf673487ae0bb25b900
title: 37단계
challengeType: 1
dashedName: step-37
---

# --description--

루프가 이제 여덟 번 실행되어야 합니다. 루프 본문 안에서 `i` 반복자의 값을 출력하고 어떤 일이 일어나는지 확인하세요.

# --hints--

`console.log()`을 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/)
```

`i`의 값을 출력하세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*i\s*\)/);
```

`i` 루프 안에서 `for`의 값을 출력하세요.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*console\.log\(\s*i\s*\);?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {

}
--fcc-editable-region--
```
