---
id: 660f061d259bbebc37461080
title: 20단계
challengeType: 1
dashedName: step-20
---

# --description--

배열 안의 값에 <dfn>인덱스</dfn>를 사용해 접근할 수 있습니다. 인덱스는 배열에서 값의 위치를 나타내는 숫자로, 첫 번째 값은 `0`부터 시작합니다.

`array[0]`와 같은 <dfn>대괄호 표기법</dfn>을 사용해 값을 접근할 수 있습니다.

`console.log`와 대괄호 표기법을 사용해 `rows` 배열의 첫 번째 값을 출력하세요.

# --hints--

코드에 `console.log()` 문장이 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

`rows` 배열에 접근해야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows/g), 2);
```

`rows` 배열에 대괄호 표기법을 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /rows\[/);
```

`rows` 배열의 첫 번째 요소에 대괄호 표기법을 사용해 접근해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*0\s*\]/);
```

`rows` 배열의 첫 번째 요소를 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*]\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];

--fcc-editable-region--
```
