---
id: 660f34626216270c682e2f7b
title: 66단계
challengeType: 1
dashedName: step-66
---

# --description--

<dfn>함수 호출</dfn>은 함수를 실제로 사용하는 방법입니다. 눈치채지 못했을 수도 있지만, 지금까지 사용해온 `.push()` 같은 메서드들은 함수 호출이었습니다.

함수는 함수 이름을 참조하고 `()`를 추가하여 호출합니다. `test` 함수를 호출하는 방법은 다음과 같습니다:

```js
test();
```

`character.repeat(i + 1)` 호출에서 `.push()`를 `padRow` 함수 호출로 바꾸세요. 아직 인수는 추가하지 마세요.

# --hints--

`i + 1` 호출에서 `push`를 사용하지 마세요.

```js
assert.notMatch(__helpers.removeJSComments(code), /repeat\(\s*i\s*\+\s*1\s*\)/);
```

`character.repeat` 호출에서 `.push()`를 사용하지 마세요.

```js
assert.notMatch(__helpers.removeJSComments(code), /push\(\s*character/);
```

`padRow` 호출에서 `.push()`를 호출해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\((.+?)?\)\)/);
```

`padRow` 호출에 인수가 없어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\(\s*?\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
