---
id: 660f039ff313dbb696b007ca
title: 19단계
challengeType: 1
dashedName: step-19
---

# --description--

배열이 값을, 즉 <dfn>요소</dfn>를 담고 있을 때, 그 값들은 쉼표로 구분됩니다. 다음은 두 개의 문자열을 담고 있는 배열입니다:

```js
let array = ["first", "second"];
```

`rows` 선언을 `"Naomi"`, `"Quincy"`, `"CamperChan"` 문자열을 가진 배열로 바꾸세요. 배열 내 값의 순서는 중요하니 그 순서를 따르세요. 문자열은 대소문자를 구분한다는 점을 기억하세요.

# --hints--

배열의 첫 번째 요소는 `"Naomi"` 문자열이어야 합니다.

```js
assert.equal(rows[0], "Naomi");
```

배열의 두 번째 요소는 `"Quincy"` 문자열이어야 합니다.

```js
assert.equal(rows[1], "Quincy");
```

배열의 세 번째 요소는 `"CamperChan"` 문자열이어야 합니다.

```js
assert.equal(rows[2], "CamperChan");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = [];
--fcc-editable-region--
```
