---
id: 660f2b6fd54ac1fc142804dd
title: 47단계
challengeType: 1
dashedName: step-47
---

# --description--

이 피라미드 형식을 만드는 논리는 복잡해질 가능성이 높으므로, 그 코드를 함수로 분리하기에 아주 좋은 시기입니다.

<dfn>함수</dfn>는 애플리케이션 전반에서 재사용할 수 있는 코드 블록입니다. 함수는 다음 구문으로 선언합니다:

```js
function name(parameter) {

}
```

`function` 키워드는 JavaScript에 `name` 변수가 함수가 될 것임을 알립니다. `parameter`는 함수가 사용될 때 전달되는 값을 나타내는 변수입니다. 함수는 원하는 만큼 매개변수를 가질 수 있습니다. `for` 반복문처럼, 중괄호 사이 공간이 <dfn>함수 본문</dfn>입니다.

`padRow` 함수를 선언하세요. 아직 매개변수 변수는 만들지 마세요. 함수 본문은 비워 두어야 합니다. 명명 관례로 카멜 케이스를 사용해야 한다는 점을 기억하세요.

# --hints--

`function` 키워드를 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /function/);
```

`padRow` 함수를 선언해야 합니다.

```js
assert.isFunction(padRow);
```

`padRow()` 함수에는 매개변수가 없어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)/);
```

`padRow()` 함수 본문은 비워 두어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
