---
id: 660f46460f9c36330ebc07d8
title: 88단계
challengeType: 1
dashedName: step-88
---

# --description--

동등 연산자는 JavaScript에서 때때로 이상한 동작을 일으킬 수 있습니다. 예를 들어, `"0" == 0`는 하나는 문자열이고 하나는 숫자임에도 true입니다.

<dfn>엄격한 동등 연산자</dfn> `===`는 두 값이 같고 같은 타입인지 확인하는 데 사용됩니다. 일반적으로 항상 이 동등 연산자를 사용해야 합니다. 엄격한 동등 연산자를 사용하면 `"0" === 0`는 false가 되는데, 값이 0으로 같을 수는 있어도 타입이 같지 않기 때문입니다.

`done == count` 조건을 엄격한 동등 연산자를 사용하도록 업데이트하세요.

# --hints--

`if` 조건에 엄격한 동등 연산자를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done == count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
