---
id: 660f4e74f7fd3f4a99ac2e50
title: 108단계
challengeType: 1
dashedName: step-108
---

# --description--

배열에는 `.shift()` 메서드도 있습니다. 이 메서드는 배열의 **첫 번째** 요소를 제거하는데, 마지막 요소를 제거하는 `.pop()`과는 다릅니다. 다음은 `.shift()` 메서드의 예입니다:

```js
const numbers = [1, 2, 3];
numbers.shift();
```

`numbers` 배열은 `[2, 3]`가 됩니다.

`numbers` 배열 바로 아래에 `shifted` 변수를 선언하고 `.shift()` 배열에서 `numbers`을 호출한 결과를 할당하세요. 다음 줄에는 `shifted` 변수를 콘솔에 출력하세요.

# --hints--

`const`을 사용해 `shifted` 변수를 선언하세요.

```js
assert.match(__helpers.removeJSComments(code), /const\s+shifted/);
```

`.shift()` 배열에서 `numbers`을 호출하세요.

```js
assert.deepEqual(numbers, [5, 2, 3]);
```

`.shift()` 호출 결과를 `shifted` 변수에 할당하세요.

```js
assert.equal(shifted, 1);
```

`shifted` 변수를 출력하세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*shifted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];


const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
