---
id: 660f505d02b2bd513a1c3468
title: 111단계
challengeType: 1
dashedName: step-111
---

# --description--

피라미드가 더 이상 거꾸로 되지 않았습니다. 이는 배열의 **끝**에 새 행을 추가하고 있기 때문입니다.

루프 본문을 업데이트하여 새 행을 배열의 시작 부분에 추가하세요.

# --hints--

`unshift`의 `rows` 메서드를 사용하세요.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /\.unshift/);
```

`padRow()` 메서드의 인수로 `.unshift()` 호출을 전달하세요.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(/);
```

`i` 호출의 첫 번째 인수로 `padRow()`를 전달하세요.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i/)
```

`count` 호출의 두 번째 인수로 `padRow()`를 전달하세요.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
