---
id: 660f383d4c772c12ff59904b
title: 70단계
challengeType: 1
dashedName: step-70
---

# --description--

전체 표현식을 인수로 전달할 수 있습니다. 함수는 그 표현식을 평가한 결과를 받습니다. 예를 들어, 아래 두 함수 호출은 같은 결과를 낼 것입니다. (연산 순서 규칙인 PEMDAS—괄호, 지수, 곱셈, 나눗셈, 덧셈, 뺄셈—가 적용된다는 점을 참고하세요):

```js
test(2 * 3 + 1);
test(7);
```

패턴을 다시 살펴보면:

```md
··#··
·###·
#####
```

`character` 값을 `2 * rowNumber - 1` 번 반복하도록 업데이트하세요.

콘솔을 다시 열어 업데이트된 결과를 확인하세요.

# --hints--

함수가 `2 * rowNumber - 1`를 `.repeat()` 호출에 전달해야 합니다.

```js
assert.equal(padRow(4, 5), " ####### ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(rowNumber) + " ".repeat(rowCount - rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
