---
id: 660f3dd626be3a1ffe27e5d1
title: 77단계
challengeType: 1
dashedName: step-77
---

# --description--

JavaScript는 여러 줄 주석도 지원합니다. 여러 줄 주석은 `/*`로 시작해서 `*/`로 끝납니다.

한 줄 주석과 달리 여러 줄 주석은 여러 줄을 감쌀 수 있습니다.

현재 `/*` 반복문 전체, 본문을 포함해서 `*/`와 `for`를 사용해 여러 줄 주석으로 만드세요.

# --hints--

여러 줄 주석은 `/*`로 시작해야 합니다.

```js
assert.match(code, /\/\*/);
```

여러 줄 주석은 `*/`로 끝나야 합니다.

```js
assert.match(code, /\*\//);
```

전체 `for` 반복문이 주석 처리되어야 합니다.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for/g), 1);
assert.notMatch(stripped, /rows\.push/);
assert.notMatch(stripped, /i <= count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
