---
id: 660f1e3f047bf4e403268713
title: 39단계
challengeType: 1
dashedName: step-39
---

# --description--

안타깝게도 이제 루프가 무엇을 하는지 볼 수 없습니다.

`let`를 사용해 `result` 변수를 선언하고 빈 문자열을 할당하세요. 빈 문자열은 `""`처럼 따옴표 사이에 아무것도 없는 것으로 나타냅니다.

# --hints--

`result` 변수를 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /result/);
```

`let`를 사용해 `result` 변수를 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /let\s+result/);
```

`result` 변수는 빈 문자열이어야 합니다.

```js
assert.equal(result, "");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

--fcc-editable-region--

--fcc-editable-region--
```
