---
id: 660f165270622fd4ec0da3f7
title: 31단계
challengeType: 1
dashedName: step-31
---

# --description--

`let` 키워드로 변수를 선언하면 다시 할당할 수 있습니다. 즉, 나중에 `character` 값을 완전히 다르게 바꿀 수 있습니다.

이 프로젝트에서는 이러한 변수 값을 변경하지 않을 것입니다. 대신 `const`를 사용해 선언해야 합니다. `const` 변수는 특별합니다.

먼저, `const` 변수는 `let` 변수처럼 다시 할당할 수 없습니다. 이 코드는 오류를 발생시킵니다:

```js
const firstName = "Naomi";
firstName = "Jessica";
```

`const` 변수는 초기화하지 않고 선언할 수도 없습니다. 이 코드는 오류를 발생시킵니다:

```js
const firstName;
```

`let` 키워드를 `const`로 바꾸세요.

# --hints--

`const`를 사용해 `character` 변수를 선언하세요.

```js
assert.match(__helpers.removeJSComments(code), /const\s+character/);
```

`const`를 사용해 `count` 변수를 선언하세요.

```js
assert.match(__helpers.removeJSComments(code), /const\s+count/);
```

`const`를 사용해 `rows` 변수를 선언하세요.

```js
assert.match(__helpers.removeJSComments(code), /const\s+rows/);
```

코드에서 `let`를 사용하지 마세요.

```js
assert.notMatch(__helpers.removeJSComments(code), /let/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
let count = 8;
let rows = [];
--fcc-editable-region--
```
