---
id: 660f4f79e2a82a4e92290f44
title: 110단계
challengeType: 1
dashedName: step-110
---

# --description--

가끔 주석 처리한 이전 코드를 다시 활성화하고 싶을 때가 있습니다. 그럴 때는 해당 코드 주변의 `/*`와 `*/`를 제거하면 됩니다. 이것을 <dfn>주석 해제</dfn>라고 합니다.

첫 번째 `for` 루프만 주석 해제하세요. 한 줄 주석과 나머지 두 개의 여러 줄 주석은 그대로 두세요.

# --hints--

첫 번째 `for` 루프를 주석 해제해야 합니다.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /for\s*\(\s*let\s+i\s*=\s*1;\s*i\s*<=\s*count;\s*i\+\+\s*\)/)
```

한 줄 주석은 제거하지 마세요.

```js
assert.match(code, /\/\//);
```

`while` 루프는 주석 해제하지 마세요.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

두 번째 `for` 루프도 주석 해제하지 마세요.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for\s*\(/g), 2)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
