---
id: 660f255022991ef34ed0ee88
title: 44단계
challengeType: 1
dashedName: step-44
---

# --description--

숫자를 출력하는 것만으로는 시각적으로 보기 좋은 피라미드를 만들 수 없습니다. 이제 `rows` 배열의 형식이 지정된 내용을 출력하고 있으니, 원래의 반복문을 업데이트할 차례입니다.

배열에 `i`을 추가하는 대신, `character` 변수의 값을 추가하세요.

# --hints--

더 이상 `i` 변수를 추가하지 않아야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

`character` 변수를 추가해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
