---
id: 660f2a70ad6225fa503e71c3
title: 46단계
challengeType: 1
dashedName: step-46
---

# --description--

조금 더 가까워졌습니다! 이 시점에서 JavaScript 같은 0부터 시작하는 인덱스 언어에서 자주 발생하는 <dfn>경계값 오류</dfn>를 만나고 있습니다.

`rows` 배열의 첫 번째 인덱스는 `0`이기 때문에 `for` 반복문을 `i = 0`부터 시작합니다. 하지만 문자열을 0번 반복하면 출력할 내용이 없습니다.

이 문제를 고치려면 `1` 호출에서 `i` 값에 `.repeat()`를 더하세요. 반복문 조건에서처럼 `i`에 다시 할당하지 마세요.

# --hints--

`1` 메서드에서 `i`에 `.repeat()`를 더해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
