---
id: 6610c2d8d67563174fcf96dc
title: 52단계
challengeType: 1
dashedName: step-52
---

# --description--

코드에 명시적으로 작성된 값이 있을 때, 예를 들어 함수 안의 `"Hello!"` 문자열처럼, 이를 <dfn>하드코딩</dfn>했다고 합니다. 함수 안에 값을 하드코딩하면 원하는 만큼 재사용하기 어려울 수 있습니다.

대신 함수에 <dfn>매개변수</dfn>를 정의할 수 있습니다. 매개변수는 함수를 호출할 때 값이 주어지는 특별한 변수로, 함수 내에서 사용되어 함수 코드의 결과를 동적으로 바꿀 수 있습니다.

함수에 매개변수를 추가하려면 소괄호 안에 변수 이름을 넣어야 합니다. 예를 들어, 이 `demo` 함수에는 `name` 매개변수가 있습니다:

```js
function demo(name) {

}
```

`name`는 유용한 매개변수처럼 보이니, `padRow` 함수에 추가하세요.

# --hints--

`padRow` 함수에는 `name` 매개변수가 있어야 합니다.

```js
assert.match(padRow.toString(), /\(name\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
