---
id: 6610b9f7619764fad5fd516d
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

변수의 값을 다른 변수에 할당할 수도 있습니다. 예를 들어:

```js
let first = "One";
let second = "Two";
second = first;
```

`second` 변수는 이제 `"One"` 값을 가지게 됩니다.

이를 직접 확인하려면 `secondCharacter`에서 `"Test"` 변수로 `character` 할당을 변경하세요.

그런 다음 콘솔을 열어 어떤 값이 기록되는지 확인하세요.

# --hints--

`"Test"` 값을 `secondCharacter` 변수에 할당하면 안 됩니다.

```js
assert.notEqual(secondCharacter, "Test");
```

`character` 변수의 값을 `secondCharacter` 변수에 할당하세요. 세미콜론을 잊지 마세요.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=\s*character;/);
```

`secondCharacter` 변수는 이제 `"World"` 값을 가져야 합니다.

```js
assert.equal(secondCharacter, "World");
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
--fcc-editable-region--
secondCharacter = "Test";
--fcc-editable-region--
console.log(secondCharacter);
```
