---
id: 6610b741b54b90f0c0fb3d58
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

변수 이름이 두 단어 이상일 때는 단어를 어떻게 대문자로 표기할지에 관한 특정 명명 관례가 있습니다. JavaScript에서는 <dfn>camel case</dfn> 관례를 사용합니다.

Camel case는 이름의 첫 단어는 모두 소문자로 작성하고, 그 다음 단어들은 모두 첫 글자를 대문자로 작성하는 것을 의미합니다. 다음은 camel case 예시입니다:

```js
let variableOne;
let secondVariable;
let yetAnotherVariable;
let thisIsAnAbsurdlyLongName;
```

camel case를 사용해 새로운 `secondCharacter` 변수를 선언하세요.

# --hints--

`secondCharacter` 변수를 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter/);
```

`let`를 사용해 `secondCharacter` 변수를 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter/);
```

`secondCharacter` 변수에 값을 할당하지 마세요. 세미콜론을 잊지 마세요.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
console.log(character);

--fcc-editable-region--
```
