---
id: 6610c16c4fa0df12c0e30675
title: 49단계
challengeType: 1
dashedName: step-49
---

# --description--

`padRow` 함수를 호출하고 있지만, 그 함수 호출 결과를 아무 데도 사용하지 않고 있습니다. JavaScript의 모든 함수는 값을 <dfn>반환</dfn>하는데, 이는 호출 결과를 다른 곳에서 사용할 수 있도록 제공한다는 뜻입니다.

`padRow` 함수를 호출한 결과를 확인하려면, `call` 변수를 선언하고 기존 `padRow` 호출 결과를 그 변수에 할당하세요.

# --hints--

`call` 변수를 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+call/);
```

`const`를 사용해 `call` 변수를 선언하세요.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call/);
```

`call`에 `padRow` 호출 결과를 할당해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call\s*=\s*padRow\(\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow() {

}
--fcc-editable-region--
padRow();
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
