---
id: 6610c424b7119919b62932f4
title: 53단계
challengeType: 1
dashedName: step-53
---

# --description--

함수가 하드코딩된 값을 반환할 필요는 없습니다. 변수에 저장된 값을 반환할 수도 있습니다. 매개변수는 함수의 특별한 변수이므로, 매개변수도 반환할 수 있습니다.

`padRow` 함수를 변경하여 `return` 매개변수를 직접 `name` 하세요.

# --hints--

`padRow` 함수는 `name` 매개변수의 값을 반환해야 합니다.

```js
assert.equal(padRow("Naomi"), "Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
