---
id: 6610c77d50636722e5b6be17
title: 60단계
challengeType: 1
dashedName: step-60
---

# --description--

`return` 문장 아래에 문자열 `"This works!"`을 콘솔에 기록하세요.

그렇게 하면 문자열 `"This works!"`가 콘솔에 표시되지 않고 `console.log("This works!")` 줄이 회색으로 표시되는 것을 볼 수 있습니다.

콘솔 로그를 복사해서 `return` 문장 위에 붙여넣으세요. 이제 문자열 `"This works!"`가 콘솔에 나타나야 합니다.

`return` 키워드에 대해 알아야 할 중요한 점은 함수에서 반환할 값을 정의하는 것뿐만 아니라 함수나 블록 문 내에서 코드 실행을 정지시킨다는 것입니다. 즉, `return` 문장 뒤의 코드는 실행되지 않습니다.

# --hints--

`padRow` 함수는 `test` 변수를 반환해야 합니다.

```js
assert.equal(padRow("Naomi"), "Testing");
```

첫 번째 `console.log`은 `return` 키워드 뒤에 와야 합니다.

```js
assert.match(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

두 번째 `console.log`은 `return` 키워드 앞에 와야 합니다.

```js
assert.match(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  
  return test;

}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
