---
id: 6610c71600966a2191d3a64a
title: 59단계
challengeType: 1
dashedName: step-59
---

# --description--

함수에서 반환된 값은 함수를 호출하여 사용합니다. 함수 호출을 반환값으로 직접 사용할 수도 있고, 반환된 값을 변수에 저장할 수도 있습니다. 이렇게 하면 함수 내부에서 생성된 지역 범위 변수에 할당된 값을 함수 외부에서 사용할 수 있습니다.

```js
function getName() {
  const name = "Camper cat";
  return name;
}

console.log(getName()); // "Camper cat"

const capturedReturnValue = getName();
console.log(capturedReturnValue); // "Camper cat"

console.log(name); // reference error
```

`"Testing"` 값을 사용하려면 `padRow` 함수에서 값을 반환하도록 `return` 문을 업데이트하여 `test` 변수만 반환하세요.


# --hints--

`padRow` 함수는 `test` 변수를 반환해야 합니다.

```js
assert.equal(padRow("Naomi"), "Testing");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";  
  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
