---
id: 6610c1d97b1671140f95cfbb
title: 50단계
challengeType: 1
dashedName: step-50
---

# --description--

이제 `call` 변수의 값을 출력하는 로그 문을 추가하세요.

# --hints--

`console.log` 호출을 추가해야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2)
```

`call` 변수를 출력해야 합니다. 세미콜론을 잊지 마세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*call\s*\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
