---
id: 6610c21b3ef82015573ffbbe
title: 51단계
challengeType: 1
dashedName: step-51
---

# --description--

`call` 변수에 값을 정의했는데도 `undefined` 값을 가지고 있습니다! 이는 `padRow` 함수가 현재 값을 반환하지 않기 때문입니다. 기본적으로 함수는 반환값으로 `undefined`를 돌려줍니다.

다른 값을 반환하려면 `return` 키워드를 사용해야 합니다. 다음은 문자열 `"Functions are cool!"`를 반환하는 함수 예제입니다:

```js
function demo() {
  return "Functions are cool!";
}
```

함수가 문자열 `return`를 반환하도록 `"Hello!"` 키워드를 사용하세요.

# --hints--

`padRow` 함수는 `return` 키워드를 사용해야 합니다.

```js
assert.match(padRow.toString(), /return/);
```

`padRow` 함수는 문자열 `"Hello!"`를 반환해야 합니다.

```js
assert.equal(padRow(), "Hello!");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
