---
id: 6610c48c4ea0891afa7c4696
title: 54단계
challengeType: 1
dashedName: step-54
---

# --description--

콘솔을 다시 열면 `padRow` 함수가 반환값을 정의했음에도 불구하고 `undefined`을 반환하는 것을 볼 수 있습니다! 이는 매개변수에 값을 **호출할 때** 전달해야 하기 때문입니다.

함수 호출에 값을 전달할 때, 그 값을 <dfn>인수</dfn>라고 합니다. 다음은 `demo` 함수를 호출하고 `"Naomi"`을 `name` 매개변수의 인수로 전달하는 예입니다.

```js
function demo(name) {
  return name;
}
demo("Naomi");
```

`name` 호출에서 `padRow` 매개변수에 자신의 이름을 인수로 전달하세요. 이름은 문자열이므로 따옴표를 사용해야 합니다.

# --hints--

`padRow()` 호출에 문자열을 전달해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*("|').+\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow();
--fcc-editable-region--
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
