---
id: 6610c87eac0f0b256d7b037e
title: 62단계
challengeType: 1
dashedName: step-62
---

# --description--

함수가 더 이상 매개변수를 사용하지 않으므로 인수를 변경해도 영향을 주지 않았습니다.

`test` 함수에서 `return` 선언과 `padRow` 문을 제거하여 함수가 다시 비어 있도록 하세요.

# --hints--

`padRow` 함수에는 `test` 변수가 없어야 합니다.

```js
assert.notMatch(padRow.toString(), /test/);
```

`padRow` 함수는 값을 반환하지 않아야 합니다.

```js
assert.notMatch(padRow.toString(), /return/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  const test = "Testing";
  return test;
}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
