---
id: 6610c105bbdacc114d6cdc44
title: 48단계
challengeType: 1
dashedName: step-48
---

# --description--

함수를 사용하려면 호출해야 합니다. <dfn>함수 호출</dfn>은 애플리케이션에 함수의 코드를 호출한 위치에서 실행하라고 지시합니다. 함수 호출의 구문은 함수 이름 다음에 소괄호가 옵니다. 예를 들어, 이 코드는 `test` 함수를 정의하고 호출합니다.

```js
function test() {

}
test();
```

`padRow` 함수를 호출하세요.

# --hints--

`padRow` 함수를 호출해야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

`padRow` 함수 본문은 비워 두세요.

```js
assert.match(padRow.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
