---
id: 6610b8017d1671f2814e8c77
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

변수를 초기화하지 않고 선언하면 <dfn>초기화되지 않은</dfn> 상태로 간주됩니다. 현재 `secondCharacter` 변수는 초기화되지 않은 상태입니다.

`console.log()`을 추가해서 `secondCharacter` 변수의 값이 무엇인지 확인하세요.

# --hints--

두 번째 `console.log`을 추가해야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

`secondCharacter` 변수를 로그로 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*secondCharacter\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

--fcc-editable-region--
```
