---
id: 6610bbed59bc2a0194d85533
title: 22단계
challengeType: 1
dashedName: step-22
---

# --description--

`rows` 배열 안의 값이 직접 변경된 것을 확인하세요. 이것을 <dfn>변이</dfn>라고 합니다. 배열에 대해 더 배우면서 언제 배열을 변이해야 하는지, 언제 변이하지 말아야 하는지 알게 될 것입니다.

다음 단계로 넘어가기 전에, 배열을 자주 사용하는 방법을 배우기에 좋은 기회입니다. 현재 코드는 `rows[2]`로 배열의 마지막 요소에 접근하고 있습니다. 하지만 마지막 요소를 얻고 싶을 때 배열에 몇 개의 요소가 있는지 모를 수도 있습니다.

배열의 `.length` 속성을 활용할 수 있습니다. 이 속성은 배열의 요소 개수를 반환합니다. 어떤 배열의 마지막 요소를 얻으려면 다음 구문을 사용할 수 있습니다:

```js
array[array.length - 1]
```

`array.length`는 배열의 요소 개수를 반환합니다. `1`를 빼면 배열의 마지막 요소 인덱스를 얻을 수 있습니다. 이 개념을 `rows` 배열에도 적용할 수 있습니다.

`rows[2]`를 업데이트하여 `rows` 배열의 마지막 요소에 동적으로 접근하세요. 위 예제를 참고하세요.

콘솔에는 아무런 변화가 없어야 합니다.

# --hints--

`.length` 배열의 `rows` 속성을 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /rows\.length/);
```

대괄호 표기법에 `rows.length`을 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length/);
```

대괄호 표기법에서 길이에서 `1`를 빼세요.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--
rows[2] = 10;
--fcc-editable-region--
console.log(rows);
```
