---
id: 6610c6541c82551f95e765ab
title: 58단계
challengeType: 1
dashedName: step-58
---

# --description--

변수는 함수 내부에서도 선언할 수 있습니다. 이러한 변수는 <dfn>지역 범위</dfn> 또는 <dfn>블록 범위</dfn>에 있다고 간주합니다. 함수 내부에서 선언된 변수는 해당 함수 내부에서만 사용할 수 있습니다. 함수 외부에서 접근하려 하면 참조 오류가 발생합니다.

이를 직접 확인하려면 `const`를 사용해 `test` 변수를 `padRow` 함수 안에 선언하세요. 값 `"Testing"`으로 초기화하세요.

그다음 함수 아래에서 `test`를 콘솔에 출력하려 시도하세요. 함수의 지역 범위 밖에서는 정의되지 않아 오류가 발생하는 것을 볼 수 있습니다. 테스트를 통과하고 계속하려면 `console.log`를 제거하세요.

# --hints--

함수 안에 `test` 변수를 선언해야 합니다.

```js
assert.match(padRow.toString(), /var\s+test/);
```

`test`를 값 `"Testing"`으로 초기화하세요. 세미콜론을 잊지 마세요.

```js
assert.match(padRow.toString(), /var\s+test\s*=\s*('|")Testing\1;/)
```

`test` 변수는 `return` 키워드 앞에 와야 합니다.

```js
const str = padRow.toString();
const testIndex = str.indexOf("test");
const returnLocation = str.indexOf("return");
const testLocation = testIndex !== -1 ? testIndex : returnLocation
assert.isBelow(testLocation, returnLocation);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {

  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
