---
id: 6610c538372aa61cc0f5b122
title: 57단계
challengeType: 1
dashedName: step-57
---

# --description--

JavaScript의 변수는 특정 <dfn>범위</dfn>에서만 사용할 수 있습니다. 즉, 변수를 선언한 위치가 코드 내에서 변수를 사용할 수 있는 위치를 결정합니다.

첫 번째 범위는 전역 범위입니다. 함수나 `for` 반복문 같은 "블록" 외부에 선언된 변수는 <dfn>전역 범위</dfn>에 속합니다. `character`, `count`, `rows` 변수는 모두 전역 범위에 있습니다.

변수가 전역 범위에 있으면 함수 정의 내에서 접근할 수 있습니다. 다음은 전역 `title` 변수를 사용하는 함수 예시입니다:

```js
const title = "Professor ";
function demo(name) {
  return title + name;
}
demo("Naomi")
```

이 예시는 `"Professor Naomi"`를 반환합니다. `padRow` 함수를 업데이트하여 `character` 변수를 `name` 매개변수 앞에 연결한 값을 반환하세요.

# --hints--

`padRow` 함수는 `character`을 `name` 앞에 연결해야 합니다.

```js
assert.match(padRow.toString(), /character\s*\+\s*name/);
```

`padRow` 함수는 `character + name` 결과를 반환해야 합니다.

```js
assert.equal(padRow("Naomi"), "#Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
