---
id: 6610c8cfe4cf4d278e35c156
title: 63단계
challengeType: 1
dashedName: step-63
---

# --description--

예상대로, 이제 함수가 다시 `undefined`을 반환합니다. `call` 변수는 더 이상 필요하지 않으니 `call` 선언과 `console.log` 변수에 대한 `call`을 제거하세요.

# --hints--

`call` 선언이 없어야 합니다.

```js
assert.notMatch(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)), /constcall/);
```

`call` 변수를 로그로 출력하지 않아야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /call/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
