---
id: 6610b8f6a98d25f4d485a94d
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

초기화되지 않은 변수의 기본값은 `undefined`입니다. 이는 아직 정의되지 않은 값을 나타내는 특별한 데이터 유형입니다.

초기화되지 않은 변수에도 값을 할당할 수 있습니다. 다음은 예시입니다:

```js
let uninitialized;
uninitialized = "assigned";
```

선언 아래에 `"Test"` 변수에 문자열 `secondCharacter`을 할당하세요. 콘솔을 열어 로그가 어떻게 바뀌었는지 확인하세요.

# --hints--

`secondCharacter`는 초기화하지 마세요. 초기화란 변수를 선언할 때 값을 할당하는 것을 의미합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+secondCharacter\s*=/);
```

`secondCharacter`에 대입 연산자를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=/);
```

`"Test"` 변수에 문자열 `secondCharacter`를 할당하세요.

```js
assert.equal(secondCharacter, "Test");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

console.log(secondCharacter);
--fcc-editable-region--
```
