---
id: 6610bf6fa14d700beed1b109
title: 87단계
challengeType: 1
dashedName: step-87
---

# --description--

<dfn>동등 연산자</dfn> `==`는 두 값이 같은지 확인하는 데 사용됩니다. 두 값을 비교하려면 `value == 8`와 같은 문장을 사용하세요.

루프 안 `done++` 아래에 `if` 문장을 추가하세요. 이 문장은 동등 연산자를 사용해 `done`가 `count`와 같은지 확인해야 합니다.

# --hints--

루프 안에 `if` 문장을 사용해야 합니다. `done++` 다음에 추가하세요.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if/);
```

`if` 문장은 조건에서 동등 연산자를 사용해 `done`와 `count`을 비교해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*==\s*count|count\s*==\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
