---
id: 6610c83b52583e245a079217
title: 61단계
challengeType: 1
dashedName: step-61
---

# --description--

이제 `call` 변수는 `"Testing"` 값을 가집니다. 하지만 함수가 더 이상 `name` 매개변수를 사용하지 않습니다.

함수 선언에서 `name` 매개변수를 제거한 다음, `"CamperChan"` 호출에서 `padRow` 문자열도 제거하세요.

또한 `console.log` 함수에서 `padRow` 둘 다 제거하세요.
# --hints--

`padRow` 함수는 `name` 매개변수를 가지면 안 됩니다.

```js
assert.notMatch(padRow.toString(), /name/);
```

`"CamperChan"` 호출에 `padRow`을 전달하면 안 됩니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /CamperChan/);
```

여전히 `padRow` 함수를 호출해야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

`console.log` 키워드 앞에 `return`이 있으면 안 됩니다.

```js
assert.notMatch(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

`console.log` 키워드 뒤에 `return`이 있으면 안 됩니다.

```js
assert.notMatch(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  console.log("This works!");
  return test;
  console.log("This works!");
}
const call = padRow("CamperChan");
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
