---
id: 664599653fcd6e97104f9261
title: 83단계
challengeType: 1
dashedName: step-83
---

# --description--

때때로 모든 `if...else if` 조건이 `false`일 때 다른 코드를 실행하고 싶을 수 있습니다. 이럴 때 `else` 블록을 추가하면 됩니다.

`else` 블록은 `if`와 `else if` 블록의 조건이 충족되지 않을 때만 평가됩니다.

여기서는 `else` 블록이 `else if` 블록에 추가되었습니다.

```js

if (condition) {
  // this code will run if condition is true
} else if (condition2) {
  // this code will run if the first condition is false
} else {
  // this code will run 
  // if the first and second conditions are false
}
```

`else` 블록을 `else if` 블록에 추가하세요. `else` 블록 안에서 문자열 `"This is the else block"`을 콘솔에 출력하세요.

콘솔에서 결과를 확인하려면 `<` 문장의 `else if`를 수동으로 `>`로 변경할 수 있습니다. 그러면 조건이 `false`가 되어 `else` 블록이 실행됩니다.

# --hints--

`else` 블록이 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /else\s*\{/);
```

`else` 블록은 문자열 `"This is the else block"`을 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)This\s+is\s+the\s+else\s+block\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
