---
id: 66458f0a05df478aa627629e
title: 82단계
challengeType: 1
dashedName: step-82
---

# --description--

텍스트가 다시 사라졌습니다! 빈 문자열은 `false`로 평가되어 <dfn>거짓 같은 값</dfn>이 됩니다. 앞으로의 프로젝트에서 참 같은 값과 거짓 같은 값에 대해 더 배우게 됩니다.

`if` 문장 외에도 JavaScript에는 <dfn>else if</dfn> 문장이 있습니다. `else if` 문장은 한 코드 블록에서 여러 조건을 확인할 수 있게 해줍니다.

다음은 `else if` 문장의 구문입니다:

```js
if (condition1) {
  // code to run if condition1 is true
} else if (condition2) {
  // code to run if condition2 is true
} else if (condition3) {
  // code to run if condition3 is true
} 
```

첫 번째 조건이 `false`이면 JavaScript는 체인의 다음 조건을 확인합니다. 두 번째 조건이 `false`이면 JavaScript는 세 번째 조건을 확인하는 식으로 계속됩니다.

`if` 문장 아래에 `else if` 문장을 추가하여 `5`이 `10`보다 작은지 확인하세요. 그런 다음 `else if` 문장 본문 안에서 문자열 `"5 is less than 10"`을 콘솔에 출력하세요.

콘솔에서 결과를 확인하세요.

# --hints--

`else if` 문장이 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

`else if` 문장은 `5`이 `10`보다 작은지 확인해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(\s*5\s*<\s*10\s*\)/);
```

문자열 `"5 is less than 10"`을 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)5\s+is\s+less\s+than\s+10\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
