---
id: 664ed9344ba3495e1f67f8ea
title: 14단계
challengeType: 1
dashedName: step-14
---

# --description--

선언과 초기화를 검토했으니, 검토를 위해 작성한 코드를 제거하세요.

`character` 변수는 제거하지 마세요.

# --hints--

`profession` 변수가 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+profession/);
```

`age` 변수가 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+age/);
```

`console.log`에 대한 `age` 문장이 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*age\s*\);/);
```

`console.log`에 대한 `profession` 문장이 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*profession\s*\);/);
```

`character` 변수는 여전히 있어야 합니다.

```js
assert.strictEqual(character, "Hello");
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--
let profession = "teacher";
let age;
console.log(age);
console.log(profession);
--fcc-editable-region--
```
