---
id: 664ed97c55a99e5ffab759dc
title: 13단계
challengeType: 1
dashedName: step-13
---

# --description--

다음 단계로 넘어가기 전에 지금까지 배운 개념을 잠시 검토하세요.

`let` 키워드를 사용해 `profession` 변수와 `age` 변수를 선언하세요. `profession`에는 문자열 `"teacher"`를 할당하되, `age`에는 값을 할당하지 마세요.

두 변수를 모두 콘솔에 출력해 결과를 확인하세요.

# --hints--

`profession` 변수를 선언해야 합니다.

```js
assert.match(code, /(?:var|let|const)\s+profession/);
```

`let`를 사용해 `profession` 변수를 선언해야 합니다.

```js
assert.match(code, /let\s+profession/);
```

문자열 `"teacher"`를 `profession` 변수에 할당해야 합니다.

```js
assert.strictEqual(profession, "teacher");
```

`age` 변수를 선언해야 합니다.

```js
assert.match(code, /(?:var|let|const)\s+age/);
```

`let`를 사용해 `age` 변수를 선언해야 합니다.

```js
assert.match(code, /let\s+age/);
```

`age` 변수에는 값을 할당하지 않아야 합니다.

```js
assert.isUndefined(age);
```

`age`를 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\(\s*age\s*\);?/);
```

`profession`를 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\(\s*profession\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--

--fcc-editable-region--
```
