---
id: 665732e62eb8b30c249b5c2c
title: 25단계
challengeType: 1
dashedName: step-25
---

# --description--

이제 다음 배열 수업 세트로 넘어갈 준비가 되었습니다.

이전 단계에서 작성한 모든 코드를 삭제하세요.

# --hints--

`cities` 변수에 대한 모든 참조를 삭제해야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /cities/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let cities = ["London", "New York", "Mumbai"];
console.log(cities);
cities[cities.length - 1] = "Mexico City";
console.log(cities);
--fcc-editable-region--
console.log(rows);
```
