---
id: 665732da4815b70bb083915e
title: 24단계
challengeType: 1
dashedName: step-24
---


# --description--

지난 몇 단계에서 배열을 다루는 방법을 모두 배웠습니다. 배운 내용을 잠시 검토하세요.

먼저 `cities` 변수를 선언하고 문자열 `"London"`, `"New York"`, `"Mumbai"`를 요소로 하는 배열로 초기화하세요. 그런 다음 그 변수를 콘솔에 출력하세요.

출력한 후에는 `cities`의 마지막 요소를 문자열 `"Mexico City"`로 바꾸고, `cities` 변수를 다시 콘솔에 출력하세요.

정확히 완료하면 콘솔에 이 출력이 보여야 합니다.

```js
[ "London", "New York", "Mumbai" ]
[ "London", "New York", "Mexico City" ]
```

# --hints--

`let`를 사용해 `cities` 변수를 선언하세요.

```js
assert.match(code, /let\s+cities/);
```

문자열 `"London"`, `"New York"`, `"Mumbai"`를 요소로 하는 배열을 `cities` 변수에 할당하세요.

```js
assert.match(code, /let\s+cities\s*=\s*\[\s*("|'|`)London\1\s*,\s*("|'|`)New York\2\s*,\s*("|'|`)Mumbai\3\s*,?\s*\]/);
```

`console.log()`를 사용해 `cities` 배열 전체를 콘솔에 출력하세요.

```js
assert.match(code, /console\.log\(\s*cities\s*\)/);
```

`cities` 배열의 마지막 요소를 문자열 `"Mexico City"`로 바꾸세요. 배열의 마지막 요소는 `array[array.length - 1]`를 사용해 접근할 수 있다는 점을 기억하세요.

```js
assert.deepEqual(cities, ["London", "New York", "Mexico City"]);
```

코드에 `console.log(cities)` 문장이 두 개 있어야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*cities\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--

--fcc-editable-region--
console.log(rows);
```
