---
id: 66643c93e05093c728abdbe9
title: 55단계
challengeType: 1
dashedName: step-55
---

# --description--

다음 단계로 넘어가기 전에 함수가 어떻게 작동하는지 잠시 검토하세요.

`addTwoNumbers`라는 이름의 함수를 선언하세요. 이 함수는 두 개의 인수를 받아서 그 두 인수의 합을 반환해야 합니다.

함수는 하드코딩된 값을 사용하면 안 됩니다. 하드코딩된 함수의 예는 다음과 같습니다:

```js
function sayName(firstName, lastName) {
  return "John Doe";
}

sayName("Camper", "Cat");
```

이 함수는 `"John Doe"`와 `firstName` 매개변수에 어떤 인수가 전달되든 항상 `lastName`을 반환하므로 `"John Doe"`는 하드코딩된 값으로 간주됩니다.

`sum` 변수를 선언하고 `addTwoNumbers`과 `5`을 인수로 하여 `10` 함수를 호출한 값을 할당하세요. `sum` 변수를 콘솔에 출력하세요.

# --hints--

`addTwoNumbers`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(addTwoNumbers);
```

`addTwoNumbers` 함수는 두 개의 매개변수를 가져야 합니다.

```js
assert.lengthOf(addTwoNumbers, 2);
```

함수는 두 매개변수의 합을 반환해야 합니다.

```js
assert.strictEqual(addTwoNumbers(5,10), 15);
```

함수는 하드코딩된 값을 반환하면 안 됩니다. 즉, 어떤 두 숫자 인수에도 작동해야 합니다.

```js
assert.strictEqual(addTwoNumbers(3, 5), 8);
```

`sum` 변수를 선언해야 합니다.

```js
assert.isDefined(sum);
```

`sum` 변수는 `15` 값을 가져야 합니다.

```js
assert.strictEqual(sum, 15);
```

`sum`에 `addTwoNumbers` 함수를 `5`과 `10` 인수로 호출한 값을 할당해야 합니다.

```js
assert.match(code, /sum\s*=\s*addTwoNumbers\s*\(/);
```

`sum` 변수를 출력해야 합니다.

```js
assert.match(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--

--fcc-editable-region--

const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
