---
id: 66643c9a9e0ad1c787f5fe2b
title: 56단계
challengeType: 1
dashedName: step-56
---

# --description--

간단한 검토가 끝났으니 `addTwoNumbers` 함수, `sum` 변수, 그리고 로그 문을 제거하세요.

# --hints--

코드에 `addTwoNumbers` 함수가 없어야 합니다.

```js
assert.notMatch(code, /addTwoNumbers/);
```

코드에 `sum` 변수가 없어야 합니다.

```js
assert.notMatch(code, /sum/);
```

`sum` 변수를 로그에 출력하지 않아야 합니다.

```js
assert.notMatch(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
function addTwoNumbers(num1, num2) {
  return num1 + num2;
}

const sum = addTwoNumbers(5, 10);
console.log(sum)

--fcc-editable-region--

const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
