---
id: 68de40f53fcaf106f51cc20f
title: 建置一個 HTML 音訊和影片播放器
challengeType: 25
dashedName: build-an-html-audio-and-video-player
demoType: onClick
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `h1` 元素作為頁面的主標題。
2. 你應該有兩個 `section` 元素。
3. 在第一個 `section` 元素內，你應該有一個 `h2` 元素作為正在播放的影片標題。
4. 在 `h2` 元素下方，你應該有一個帶有 `controls` 和 `width` 屬性的 `video` 元素。`width` 屬性應該設定為 `640`。
5. 在 `video` 元素內，你應該有一個 `source` 元素，其 `src` 屬性指向一個影片檔案，並且有一個 `type` 屬性。
   - 你可以使用 `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`。
6. 在第二個 `section` 元素內，你應該有一個 `h2` 元素作為正在播放歌曲的標題。
7. 在 `h2` 元素下方，你應該有一個帶有 `controls` 和 `loop` 屬性的 `audio` 元素，以及一個指向音訊檔案的 `src` 屬性。
   - 你可以使用 `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`。
   - 或 `https://cdn.freecodecamp.org/curriculum/js-music-player/we-are-going-to-make-it.mp3`。

# --hints--

你應該有一個 `h1` 元素作為頁面的主標題。

```js
const h1 = document.querySelector('h1');
assert.exists(h1);
```

你的 `h1` 不應該是空的。

```js
const h1 = document.querySelector('h1');
assert.isNotEmpty(h1.textContent.trim());
```

你應該只擁有一個 `h1` 元素

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你應該有兩個 `section` 元素。

```js
assert.lengthOf(document.querySelectorAll('section'), 2);

```

在第一個 `section` 元素內，你應該有一個 `h2` 元素作為正在播放的影片標題。

```js
const section = document.querySelector('section');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

在 `h2` 元素下方，你應該有一個 `video` 元素。

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
assert.exists(video);
```

`video` 元素應該有一個 `controls` 屬性。

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const controlsAttr = video.controls;
assert.isTrue(controlsAttr);
```

`video` 元素應該設定 `width` 屬性為 `640`。

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const widthAttr = video.getAttribute('width');
assert.equal(widthAttr, '640');
```

在 `video` 元素內，你應該有一個 `source` 元素。

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
assert.exists(source);
```

`source` 元素應該有一個指向影片檔案的 `src` 屬性。

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const srcAttr = source.getAttribute('src');
assert.match(srcAttr, /\.(mp4|webm|ogg|avi|mov)$/i);
```

`source` 元素應該具有 `type` 屬性

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
assert.isNotEmpty(typeAttr);
```

`type` 屬性應該與 `src` 屬性的副檔名的媒體型別相符。

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
const srcAttr = source.getAttribute('src');
const srcFileExt = srcAttr.match(/[^\.]+$/g)[0];
const fileExtToMediaTypes = {
  mp4: ['video/mp4'],
  webm: ['video/webm'],
  ogg: ['video/ogg'],
  avi: ['video/avi', 'video/vnd.avi', 'video/msvideo', 'video/x-msvideo'],
  mov: ['video/quicktime']
};
const srcMediaTypes = fileExtToMediaTypes[srcFileExt];
assert.oneOf(typeAttr, srcMediaTypes);
```

在第二個 `section` 元素內，你應該有一個 `h2` 元素作為正在播放歌曲的標題。

```js
const section = document.querySelector('section:nth-of-type(2)');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

你所有的 `h2` 元素都應該包含一些文字。

```js
const h2s = document.querySelectorAll('h2');
assert.isNotEmpty(h2s);
h2s.forEach((h2) => assert.isNotEmpty(h2.textContent.trim()));
```

在 `h2` 元素下方，你應該有一個 `audio` 元素。

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
assert.exists(audio);
```

`audio` 元素應該有一個 `controls` 屬性。

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const controlsAttr = audio.controls;
assert.isTrue(controlsAttr);
```

`audio` 元素應該有一個 `loop` 屬性。

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const loopAttr = audio.loop;
assert.isTrue(loopAttr);
```

`audio` 元素應該有一個指向音訊檔案的 `src` 屬性。

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const srcAttr = audio.getAttribute('src');
assert.match(srcAttr, /\.(mp3|wav|ogg|aac|flac)$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>
  <main>
    <h1>HTML Audio and Video Lab</h1>

    <section>
      <h2>What is the map method and how does it work?</h2>
      <video controls width="640">
        <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4" />
      </video>
    </section>

    <section>
      <h2>Sailing Away</h2>
      <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" controls loop>
      </audio>
    </section>

  </main>
</body>
</html>
```
