---
id: 5cc0bd7a49b71cb96132e54c
title: 建造一個數字範圍產生器
challengeType: 26
dashedName: build-a-range-of-numbers-generator
---

# --description--

在本實驗中，你將建置一個 `rangeOfNumbers` 函式（程式），用來產生指定範圍內的數字陣列。

範例：

- `rangeOfNumbers(3, 9)` 傳回 `[3, 4, 5, 6, 7, 8, 9]`
- `rangeOfNumbers(5, 5)` 傳回 `[5]`

需求：

- 使用遞迴（該函式必須呼叫自身）
- 不使用任何迴圈或陣列方法（`for`、`while`、`forEach`、`map`、`filter`、`reduce`）

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該創建一個名為 `rangeOfNumbers` 的函式（程式），該函式接受兩個參數：`startNum` 和 `endNum`。
1. 該函式（程式）應傳回一個整數陣列，該陣列以 `startNum` 參數表述的數字開始，以 `endNum` 參數表述的數字結束（包含該數字）。
1. `startNum` 永遠會小於或相等於 `endNum`。
1. 你的函式必須透過呼叫自身來使用遞迴。它不應該使用任何迴圈語法（`for`、`while`，或是像 `forEach`、`map`、`filter` 或 `reduce` 這類的高階函式）。
1. 該函式（程式）應該控制代碼處理 `startNum` 與 `endNum` 相等的基本情況，並傳回只包含該單一數字的陣列。
1. 對於遞迴情況，該函式應該以修改過的參數呼叫自身來建造陣列，然後將當前數字添加到結果中。
1. 該函式（程式）不應依賴全域變數來快取或建置陣列。

# --hints--

你應該有一個名為 `rangeOfNumbers` 的函式（程式）。

```js
assert.isFunction(rangeOfNumbers);
```

`rangeOfNumbers` 函式應該接受兩個參數。

```js
assert.match(rangeOfNumbers.toString(), /rangeOfNumbers\s*\(\s*\w+\s*,\s*\w+\s*\)/);
```

你的函式應該傳回一個陣列。

```js
assert(Array.isArray(rangeOfNumbers(5, 10)));
```

你的程式碼不應該使用任何迴圈語法（`for` 或 `while` 或較高階的函式（程式）如 `forEach`、`map`、`filter` 或 `reduce`）。

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

`rangeOfNumbers` 應該使用遞迴（呼叫自身）來解決此挑戰。

```js
assert(
  rangeOfNumbers.toString().match(/rangeOfNumbers\s*\(.+\)/)
);
```

`rangeOfNumbers(1, 5)` 應該傳回 `[1, 2, 3, 4, 5]`。

```js
assert.deepStrictEqual(rangeOfNumbers(1, 5), [1, 2, 3, 4, 5]);
```

`rangeOfNumbers(6, 9)` 應該傳回 `[6, 7, 8, 9]`。

```js
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

`rangeOfNumbers(4, 4)` 應該傳回 `[4]`。

```js
assert.deepStrictEqual(rangeOfNumbers(4, 4), [4]);
```

`rangeOfNumbers(10, 15)` 應該傳回 `[10, 11, 12, 13, 14, 15]`。

```js
assert.deepStrictEqual(rangeOfNumbers(10, 15), [10, 11, 12, 13, 14, 15]);
```

`rangeOfNumbers(2, 8)` 應該傳回 `[2, 3, 4, 5, 6, 7, 8]`。

```js
assert.deepStrictEqual(rangeOfNumbers(2, 8), [2, 3, 4, 5, 6, 7, 8]);
```

不應該使用全域變數來快取陣列。

```js
rangeOfNumbers(1, 3);
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function rangeOfNumbers(startNum, endNum) {
  if (endNum - startNum === 0) {
    return [startNum];
  } else {
    const numbers = rangeOfNumbers(startNum, endNum - 1);
    numbers.push(endNum);
    return numbers;
  }
}
```
