---
id: a202eed8fc186c8434cb6d61
title: Costruisci un invertitore di stringhe
challengeType: 26
dashedName: build-a-string-inverter
---

# --description--

In questo laboratorio costruirai un semplice invertitore di stringhe che inverte i caratteri di una stringa data.

Per esempio, `"hello"` dovrebbe diventare `"olleh"`.

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti creare una funzione chiamata `reverseString` che prende una stringa come argomento.
2. La funzione dovrebbe restituire la stringa invertita.

# --hints--

Dovresti avere una funzione chiamata `reverseString`.

```js
assert.isFunction(reverseString);
```

`reverseString` dovrebbe prendere una stringa come argomento.

```js
assert.match(reverseString.toString(), /\s*function\s+reverseString\s*\(\s*\w+\s*\)/);
```

`reverseString("hello")` dovrebbe restituire una stringa.

```js
assert.isString(reverseString('hello'));
```

`reverseString("hello")` dovrebbe restituire la stringa `olleh`.

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

`reverseString("Howdy")` dovrebbe restituire la stringa `ydwoH`.

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

`reverseString("Greetings from Earth")` dovrebbe restituire la stringa `htraE morf sgniteerG`.

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
