---
id: 67ef8f695db52583424bd118
title: Спроєктуйте сторінку з оглядом фільму
challengeType: 25
dashedName: design-a-movie-review-page
demoType: onClick
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історія користувача:**

1. Додайте елемент `main`.
2. Усередині елемента `main` додайте елемент `h1` для назви фільму.
3. Під елементом `h1` додайте елемент `img`, який відображає обкладинку фільму. Ваш елемент `img` повинен мати описовий текст у атрибуті `alt`, який описує зображення. Ви можете використати таке зображення, якщо бажаєте: `https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png`.
4. Додайте елемент `p`, який містить короткий опис фільму.
5. Додайте ще один елемент `p`, щоб показати рейтинг фільму. Усередині нього розмістіть такі елементи у вказаному порядку:
   - Елемент `b` з текстом `Movie Rating:`.
   - Елемент `span` з атрибутом `aria-hidden`, встановленим у `true`, який містить візуальне представлення рейтингу у вигляді зірок `⭐⭐⭐⭐⭐⭐⭐⭐⭐☆`.
   - Числове значення, що представляє рейтинг, у круглих дужках (наприклад, `9.2/10`) після span.
6. Додайте елемент `h2` з текстом `Cast Members`.
7. Додайте елемент `ul`.
8. Усередині елемента `ul` додайте кілька елементів `li`, кожен з яких містить елемент `b` з іменем актора, після якого йде відповідне ім'я персонажа, перед яким стоїть текст `as`. (наприклад, `James Holloway as Ethan Carter`).

# --hints--

У вас має бути `main` елемент.

```js
assert.isNotNull(document.querySelector("main"));
```

Всередині елемента `main` має бути елемент `h1`, що містить назву фільму.

```js
const h1 = document.querySelector("main > h1");
assert.isAbove(h1?.innerText.trim().length, 0);
```

Під елементом `h1` має бути елемент `img` з описовим атрибутом `alt`.
  
```js
const img = document.querySelector("main > h1 + img");
assert.isNotNull(img);
assert.isNotEmpty(img.getAttribute("alt"));
```

Після зображення має бути елемент `p`, що містить опис фільму.

```js
const description = document.querySelector("main img + p");
assert.isNotNull(description);
assert.isNotEmpty(description.innerText.trim());
```

Має бути ще один елемент `p`, який містить рейтинг.

```js
const paragraphs = document.querySelectorAll("main p");
assert.isAtLeast(paragraphs.length, 2);
```

Всередині другого `p` елемента має бути `b` елемент з текстом `Movie Rating: `.

```js
const boldEl = document.querySelector("main p:nth-of-type(2) b");
assert.match(boldEl?.innerText.trim(), /^Movie Rating:?$/);
```

У абзаці з рейтингом має бути елемент `span`.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.isNotNull(spanEl);
```

Елемент `span` у абзаці з рейтингом має містити десять зірок, заповнених (⭐) або порожніх (☆), а після `span` — числове значення в круглих дужках.

```js  
const spanEl = document.querySelector("main p:nth-of-type(2) span");
const fullRatingText = document.querySelector("main p:nth-of-type(2)")?.textContent.replace(/\s+/g, ' ').trim();

assert.match(spanEl?.innerText.trim(), /^[⭐☆]{10}$/);
assert.match(fullRatingText, /Movie Rating:?\s*[⭐☆]{10}\s*\(\s*\d+(?:\.\d+)?\s*\/\s*10\s*\)$/);
```

Елемент `span` у абзаці з рейтингом має атрибут `aria-hidden`, встановлений у `true`.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.strictEqual(spanEl?.getAttribute("aria-hidden"), "true");
```

У вас має бути елемент `h2`.

```js
assert.exists(document.querySelector("h2"));
```

Ваш елемент `h2` має йти після другого елемента `p`.

```js
const allParagraphs = document.querySelectorAll("p");
assert.isAtLeast(allParagraphs.length, 2);

const h2 = document.querySelector("h2");
const main = document.querySelector("main");

assert.exists(h2);
assert.exists(main);

const elements = Array.from(main.querySelectorAll("*"));
const lastPIndex = elements.indexOf(allParagraphs[1]);
const h2Index = elements.indexOf(h2);

assert.isAbove(h2Index, lastPIndex);
```

Ваш елемент `h2` має містити текст `Cast Members`.

```js
const h2 = document.querySelector("h2");
assert.equal(h2?.textContent.trim(), "Cast Members");
```

Після `ul` має бути елемент `h2`.

```js
const ul = document.querySelector("main h2 + ul");
assert.isNotNull(ul);
```

`ul` має містити кілька елементів `li`.

```js
const liItems = document.querySelectorAll("main ul li");
assert.isAbove(liItems.length, 1);
```

Кожен `li` має містити `b` елемент з ім’ям актора, за яким іде ім’я відповідного персонажа, перед яким стоїть текст `as`. (наприклад, `James Holloway as Ethan Carter`).

```js
const liItems = document.querySelectorAll("main ul li");  
assert.isNotEmpty(liItems);
liItems.forEach(li => {
  const bold = li.querySelector("b");
  assert.isNotNull(bold);
  const fullText = li.textContent.trim();
  const actorName = bold.textContent.trim();
  assert.isTrue(fullText.startsWith(actorName));
  const characterPart = fullText.slice(actorName.length).trim();
  assert.match(characterPart, /^as\s.+/);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
 <html lang="en">
 <head>
     <meta charset="UTF-8">
     <meta name="viewport" content="width=device-width, initial-scale=1.0">
     <title>Movie Review</title>
 </head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Movie Review</title>
  </head>
  <body>
    <main class="review">
      <h1>Movie Review: Rise Beyond</h1>
      <img
        src="https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png"
        alt="The Rise Beyond movie poster, starring Ethan Carter. A portrait of a young boy standing on a stage with his back toward the audience, holding a guitar at his side, staring at a backdrop of a cityscape with the sun rising behind its tall buildings."
      />
      <p>
        Rise Beyond is an uplifting drama that tells the story of Ethan Carter,
        a young man from a small town who dreams of becoming a world-class
        musician. Facing financial struggles and self-doubt, Ethan embarks on a
        journey of resilience, guided by his mentor Professor Adams and
        childhood friend Lena Mitchell. Through setbacks, failures, and moments
        of triumph, Rise Beyond reminds us that our greatest potential lies
        beyond our fears.
      </p>
      <p>
        <b>Movie Rating:</b>
        <span aria-hidden="true">⭐⭐⭐⭐⭐⭐⭐⭐⭐☆</span> (9.2/10)
      </p>

      <h2>Cast Members</h2>
      <ul>
        <li><b>James Holloway</b> as Ethan Carter</li>
        <li><b>Olivia Sterling</b> as Lena Mitchell</li>
        <li><b>William Lancaster</b> as Professor Adams</li>
        <li>
          <b>Maria Collins</b> as Evelyn Carter (Ethan’s Mother)
        </li>
      </ul>
    </main>
  </body>
</html>
```
