---
id: 686386e76197458b7dfd52ec
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Abaixo do `h2`, crie um elemento `p` com o texto `$0`. Então, dentro desse elemento `p` e após o texto, crie um elemento `span` com o texto `/month`.

# --hints--

Você deve criar um elemento `p` com o texto `$0`.

```js
const p = document.querySelector("p");
const pTextBeforeSpan = p.childNodes[0].textContent.trim();

assert.equal(pTextBeforeSpan, "$0");
```

Você deve criar um elemento `span` com o texto `/month` dentro do elemento `p`.

```js
const spanText = document.querySelector("p > span")?.textContent;
assert.equal(spanText, "/month");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            --fcc-editable-region--
            
            --fcc-editable-region--
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
