---
id: 686386e76197458b7dfd52ef
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Agora você vai criar um botão CTA.Um Call to Action (CTA) é um botão ou link projetado para incentivar o usuário a realizar uma ação específica, neste caso, começar a ouvir.

Fora do `div` mais interno, crie um elemento anchor com um `href` de `#` e o texto `Start listening`.

Isso conclui o HTML para o primeiro cartão de preços.

# --hints--

Você deve criar um elemento de âncora.

```js
assert.exists(document.querySelector("a"))
```

Seu elemento de âncora deve ter um `href` definido como `#`.

```js
const link = document.querySelector("a")
assert.equal(link?.getAttribute("href"), "#")
```

Seu elemento de âncora deve ter o texto `Start listening`.

```js
assert.equal(document.querySelector("a")?.textContent, "Start listening")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          --fcc-editable-region--
          
          --fcc-editable-region-- 
        </div>
      </div>
    </main>
  </body>
</html>
```
