---
id: 686386e76197458b7dfd52ee
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Dentro do `ul`, crie dois itens de lista.Dentro do primeiro item da lista, crie um elemento `span` e adicione a entidade HTML `&#10003;` como seu texto.Logo após esse `span` mas dentro do item da lista, adicione o texto `Ad-supported streaming`.

Dentro do segundo item da lista, crie um elemento `span` e adicione a entidade `&#10003;` como seu texto.Logo após aquele `span`, mas dentro do item da lista, adicione o texto `Curated playlists`.

# --hints--

Você deve ter dois itens de lista dentro do seu elemento `ul`.

```js
assert.lengthOf(document.querySelectorAll("ul > li"), 2)
```

Seu primeiro elemento `li` deve ter um elemento `span` com a entidade de caractere HTML `&#10003;` dentro dele.

```js
const listItems = document.querySelectorAll("ul > li")
const firstSpan = listItems[0].querySelector("span")

assert.equal(firstSpan?.textContent, "✓")
```

Seu primeiro elemento `li` deve ter o texto `Ad-supported streaming` logo após o elemento `span`.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[0].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Ad-supported streaming");
```

Seu segundo elemento `li` deve ter um elemento `span` com a entidade de caractere HTML `&#10003;` dentro dele.

```js
const listItems = document.querySelectorAll("ul > li")
const secondSpan = listItems[1].querySelector("span")

assert.equal(secondSpan?.textContent, "✓")
```

Seu segundo elemento `li` deve ter o texto `Curated playlists` logo após o elemento `span`.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[1].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Curated playlists");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>Start exploring millions of songs with basic features and ads.</p>
            <ul>
            --fcc-editable-region--
              
            --fcc-editable-region--
            </ul>
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
