---
id: 686383abac6c625a84458ea9
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você aprenderá mais sobre Tailwind CSS usando suas classes utilitárias de grid para construir um componente de preços.

O boilerplate básico de HTML foi fornecido para você.Comece criando um elemento `h1` com o texto `Choose your listening plan`.

# --hints--

Você deve criar um elemento `h1`.

```js
assert.exists(document.querySelector("h1"));
```

Seu elemento `h1` deve ter o texto `Choose your listening plan`.

```js
assert.equal(document.querySelector("h1")?.textContent, "Choose your listening plan");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </main>
  </body>
</html>
```
