---
id: 686386e76197458b7dfd52eb
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Crie um elemento `div` abaixo do `h1`. Dentro desse `div`, crie outro elemento `div` e aninhe outro elemento `div` dentro dele.

E dentro do `div` mais interno, crie um elemento `h2` com o texto `Listener`. Isto representará o primeiro plano de preços.

# --hints--

Você deve criar um elemento `div`.

```js
assert.exists(document.querySelector("div"));
```

Você deve criar outro elemento `div` dentro do `div` externo.

```js
assert.exists(document.querySelector("div > div"));
```

Você deve criar outro elemento `div` dentro do `div` interno.

```js
assert.exists(document.querySelector("div > div > div"));
```

Você deve criar um `h2` com o texto `Listener` dentro do elemento `div` mais interno.

```js
assert.equal(document.querySelector("div > div > div > h2")?.textContent, "Listener");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
