---
id: a202eed8fc186c8434cb6d61
title: Побудувати інвертор рядка
challengeType: 26
dashedName: build-a-string-inverter
---

# --description--

У цій лабораторній роботі ви створите простий інвертор рядка, який перевертає символи заданого рядка.

Наприклад, `"hello"` має стати `"olleh"`.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви маєте створити функцію з назвою `reverseString`, яка приймає рядок як аргумент.
2. Функція має повертати перевернутий рядок.

# --hints--

Ви маєте мати функцію з назвою `reverseString`.

```js
assert.isFunction(reverseString);
```

`reverseString` має приймати рядок як аргумент.

```js
assert.match(reverseString.toString(), /\s*function\s+reverseString\s*\(\s*\w+\s*\)/);
```

`reverseString("hello")` має повертати рядок.

```js
assert.isString(reverseString('hello'));
```

`reverseString("hello")` має повертати рядок `olleh`.

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

`reverseString("Howdy")` має повертати рядок `ydwoH`.

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

`reverseString("Greetings from Earth")` має повертати рядок `htraE morf sgniteerG`.

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
