---
id: 673c91f0b934834bc4a3ecc2
title: Erstellen eines fCC-Forum-Ranglisten-Boards
challengeType: 25
dashedName: build-an-fcc-forum-leaderboard
demoType: onClick
---

# --description--

In diesem Labor erstellen Sie ein freeCodeCamp-Forum-Ranglisten-Board, das die neuesten Themen, Benutzer und Antworten von [freeCodeCamp forum](https://forum.freecodecamp.org/) anzeigt. Das HTML, CSS und ein Teil des JS wurden für Sie bereitgestellt. Erkunden Sie diese gerne.

Der API-Endpunkt liefert JSON in folgender Form:

```json
{
  "users": [
    {
      "id": 6,
      "username": "QuincyLarson",
      "name": "Quincy Larson",
      "avatar_template": "/user_avatar/QuincyLarson_{size}.png"
    },
    {
      "id": 576147,
      "username": "JOY-OKORO",
      "name": "Joy Okoro",
      "avatar_template": "/user_avatar/JOY-OKORO_{size}.png"
    }
  ],
  "topic_list": {
    "topics": [
      {
        "id": 684569,
        "title": "The freeCodeCamp Podcast is back – now with video",
        "slug": "the-freecodecamp-podcast-is-back-now-with-video",
        "posts_count": 8,
        "views": 542,
        "bumped_at": "2024-04-15T16:01:26.403Z",
        "category_id": 1,
        "posters": [
          { "user_id": 6 },
          { "user_id": 576147 }
        ]
      }
    ]
  }
}
```

Beachten Sie, dass `avatar_template`-Werte relative Pfade sind: Sie beginnen mit `/` und enthalten einen `{size}`-Platzhalter, der durch eine Pixelgröße ersetzt werden sollte. Die Konstante `avatarUrl` enthält die Basis-URL, die vorangestellt werden muss, wenn der Pfad relativ ist.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und lassen Sie alle Tests erfolgreich durchlaufen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion namens `timeAgo` haben, die einen Zeitstempel im ISO-8601-Format als Argument annimmt.
1. Die Funktion `timeAgo` sollte die Zeitdifferenz zwischen der als Argument übergebenen Zeit und der aktuellen Zeit berechnen und zurückgeben:
   - `xm ago` (`x` steht für Minuten), wenn die verstrichene Zeit in Minuten weniger als `60` beträgt.
   - `xh ago` (`x` steht für Stunden), wenn die verstrichene Zeit in Stunden weniger als `24` beträgt.
   - `xd ago` (`x` steht für Tage) sonst.
1. Sie sollten eine Funktion namens `viewCount` haben, die die Anzahl der Ansichten eines Beitrags als Argument annimmt.
1. Wenn der Wert der Ansichten, der als Argument übergeben wird, größer oder gleich `1000` ist, sollte die Funktion `viewCount` einen String zurückgeben, der den Wert der Ansichten geteilt durch `1000`, abgerundet auf die nächste ganze Zahl, und den Buchstaben `k` angehängt enthält. Andernfalls sollte sie den Wert der Ansichten zurückgeben.
1. Sie sollten eine Funktion namens `forumCategory` haben, die die ID einer ausgewählten Kategorie als Argument annimmt.
1. Die Funktion `forumCategory` sollte überprüfen, ob die ausgewählte Kategorie-ID eine Eigenschaft des Objekts `allCategories` ist, und einen String zurückgeben, der ein Ankerelement enthält mit:
   - dem Text des Schlüssels `category` der ausgewählten Kategorie.
   - einer Klasse `category`, gefolgt von der Eigenschaft `className` der ausgewählten Kategorie.
   - einem `href`, gebildet durch Anhängen von `<className>/<id>` an `forumCategoryUrl` (Hinweis: `forumCategoryUrl` endet bereits mit `/`, daher fügen Sie keinen zusätzlichen `/`-Trennzeichen hinzu), wobei `<className>` die Eigenschaft `className` der ausgewählten Kategorie ist und `id` das Argument ist, das an `forumCategory` übergeben wird.
1. Wenn das Objekt `allCategories` die ausgewählte Kategorie-ID nicht als Eigenschaft besitzt, sollte `category` als `General` und `className` als `general` angegeben werden.
1. Sie sollten eine Funktion namens `avatars` haben, die zwei Arrays repräsentierend für Poster und Benutzer als Argumente annimmt.
1. Die Funktion `avatars` sollte einen String zurückgeben, der durch das Verbinden von `img`-Elementen gebildet wird, eines für jeden `user_id` im Array `posters`. Finden Sie die `img`-URL, indem Sie die Eigenschaft `user_id` im Array `posters` nachschlagen und die passende Eigenschaft `id` im Array `users` finden.
1. Die Funktion `avatars` sollte die Größe jedes Avatars festlegen, indem sie auf die Eigenschaft `avatar_template` des passenden Benutzerobjekts zugreift (gefunden im Array `users` durch Vergleich seiner `id` mit `user_id` in `posters`) und den `{size}`-Platzhalter im URL-String durch `30` ersetzt.
1. Jedes Bildelement sollte einen Alt-Text mit dem Wert der Eigenschaft `name` des Posters haben.
1. Jedes `img`-Element sollte ein `src`-Attribut mit dem Wert `avatar_template` des passenden Benutzers haben. Wenn `avatar_template` mit `/` beginnt, fügen Sie `avatarUrl` direkt davor ein.
1. Sie sollten eine Funktion namens `showLatestPosts` haben, die einen einzelnen Parameter annimmt.
1. Die Funktion `showLatestPosts` sollte die Eigenschaften `users` und `topic_list` aus dem als Argument übergebenen Objekt extrahieren. Außerdem sollte sie die folgenden Eigenschaften der Objekte aus dem Array `topics`, das in `topic_list` enthalten ist, verarbeiten:
   - `id`: die ID des Beitrags
   - `title`: der Titel des Beitrags
   - `views`: die Anzahl der Ansichten des Beitrags
   - `posts_count`: die Anzahl der Antworten zum Thema
   - `slug`: der Slug des Beitrags
   - `posters`: die Poster für dieses Thema
   - `category_id`: eine Ganzzahl, die die Kategorie-ID für den Beitrag angibt
   - `bumped_at`: ein Zeitstempel im ISO-8601-Format
1. Die Funktion `showLatestPosts` sollte das innere HTML von `#posts-container` auf einen String setzen, der durch das Verbinden von `tr`-Elementen gebildet wird, eines für jeden Eintrag in `topics`.
1. Jedes `tr`-Element sollte fünf `td`-Elemente enthalten:
   - ein `td`, das zwei Ankerelemente enthält, eines mit der Klasse `post-title`, einem `href` von `<forumTopicUrl><slug>/<id>`, einem Ankertext von `<title>` und eines, das durch Aufrufen von `forumCategory` mit `category_id` erhalten wird.
   - ein `td`, das ein `div`-Element mit der Klasse `avatar-container` enthält, das die Bilder enthält, die von der Funktion `avatars` zurückgegeben werden, die mit `posters` und `users` als Argumente aufgerufen wird.
   - ein `td`, das die Anzahl der Antworten auf den Beitrag enthält. _Hinweis:_ Verwenden Sie `posts_count - 1`.
   - ein `td`, das die Anzahl der Ansichten des Beitrags enthält.
   - ein `td`, das die seit der letzten Aktivität verstrichene Zeit enthält.
1. Sie sollten eine asynchrone Funktion namens `fetchData` haben.
1. Die Funktion `fetchData` sollte Daten von `forumLatest` anfordern und `showLatestPosts` mit der als JSON geparsten Antwort aufrufen.
1. Wenn beim Abrufen der Daten ein Fehler auftritt, sollte die Funktion `fetchData` den Fehler in der Konsole protokollieren. Verwenden Sie dafür speziell `console.log`.

# --hints--

Sie sollten eine Funktion namens `timeAgo` haben, die ein einzelnes Argument annimmt.

```js
assert.isFunction(timeAgo);
assert.lengthOf(timeAgo, 1);
```

Wenn die Zeitdifferenz zwischen der als Argument übergebenen Zeit und der aktuellen Zeit `50` Minuten beträgt, sollte `timeAgo` `50m ago` zurückgeben.

```js
const generateTime = () => {
  const currentTime = new Date();
  return new Date(currentTime - 1000 * 60 * 50).toISOString();
};
const expected = '50m ago';
const actual = timeAgo(generateTime());
assert.equal(actual, expected);
```

Wenn die Zeitdifferenz zwischen der als Argument übergebenen Zeit und der aktuellen Zeit `60` Minuten beträgt, sollte `timeAgo` `1h ago` zurückgeben.

```js
const generateTime = () => {
  const currentTime = new Date();
  return new Date(currentTime - 1000 * 60 * 60).toISOString();
};
const expected = '1h ago';
const actual = timeAgo(generateTime());
assert.equal(actual, expected);
```

Wenn die Zeitdifferenz zwischen der als Argument übergebenen Zeit und der aktuellen Zeit `115` Minuten beträgt, sollte `timeAgo` `1h ago` zurückgeben.

```js
const generateTime = () => {
  const currentTime = new Date();
  return new Date(currentTime - 1000 * (60 * 115)).toISOString();
};
const expected = '1h ago';
const actual = timeAgo(generateTime());
assert.equal(actual, expected);
```

Wenn die Zeitdifferenz zwischen der als Argument übergebenen Zeit und der aktuellen Zeit `15` Stunden beträgt, sollte `timeAgo` `15h ago` zurückgeben.

```js
const generateTime = () => {
  const currentTime = new Date();
  return new Date(currentTime - 1000 * 60 * 60 * 15).toISOString();
};
const expected = '15h ago';
const actual = timeAgo(generateTime());
assert.equal(actual, expected);
```

Wenn die Zeitdifferenz zwischen der als Argument übergebenen Zeit und der aktuellen Zeit `24` Stunden beträgt, sollte `timeAgo` `1d ago` zurückgeben.

```js
const generateTime = () => {
  const currentTime = new Date();
  return new Date(currentTime - 1000 * 60 * 60 * 24).toISOString();
};
const expected = '1d ago';
const actual = timeAgo(generateTime());
assert.equal(actual, expected);
```

Wenn die Zeitdifferenz zwischen der als Argument übergebenen Zeit und der aktuellen Zeit `46` Stunden beträgt, sollte `timeAgo` `1d ago` zurückgeben.

```js
const generateTime = () => {
  const currentTime = new Date();
  return new Date(currentTime - 1000 * 60 * 60 * 46).toISOString();
};
const expected = '1d ago';
const actual = timeAgo(generateTime());
assert.equal(actual, expected);
```

Wenn die Zeitdifferenz zwischen der als Argument übergebenen Zeit und der aktuellen Zeit `3` Tage beträgt, sollte `timeAgo` `3d ago` zurückgeben.

```js
const generateTime = () => {
  const currentTime = new Date();
  return new Date(currentTime - 1000 * 60 * 60 * 24 * 3).toISOString();
};
const expected = '3d ago';
const actual = timeAgo(generateTime());
assert.equal(actual, expected);
```

Sie sollten eine Funktion namens `viewCount` haben, die ein einzelnes Argument annimmt.

```js
assert.isFunction(viewCount);
assert.lengthOf(viewCount, 1);
```

`viewCount(597)` sollte `597` zurückgeben.

```js
assert.strictEqual(597, viewCount(597));
```

`viewCount(1000)` sollte `1k` zurückgeben.

```js
assert.equal('1k', viewCount(1000));
```

`viewCount(2730)` sollte `2k` zurückgeben.

```js
assert.equal('2k', viewCount(2730));
```

Sie sollten eine Funktion namens `forumCategory` haben, die ein einzelnes Argument annimmt.

```js
assert.isFunction(forumCategory);
assert.lengthOf(forumCategory, 1);
```

`forumCategory(299)` sollte einen String zurückgeben, der ein Ankerelement mit dem Text `Career Advice` enthält.

```js
let actual = forumCategory(299);
assert.match(actual.trim(), /^<\s*a.+?>\s*Career Advice\s*<\/a>$/);
// prevent hardcoding

actual = forumCategory(409);
assert.match(actual.trim(), /^<\s*a.+?>\s*Project Feedback\s*<\/a>$/);
```

`forumCategory(299)` sollte einen String zurückgeben, der ein Ankerelement mit `href="https://forum.freecodecamp.org/c/career/299"` enthält.

```js
let actual = forumCategory(299);
assert.match(
  actual,
  /href=("|')https:\/\/forum\.freecodecamp\.org\/c\/career\/299\1/
);

// prevent hardcoding
actual = forumCategory(409);
assert.match(
  actual,
  /href=("|')https:\/\/forum\.freecodecamp\.org\/c\/feedback\/409\1/
);
```

`forumCategory(299)` sollte einen String zurückgeben, der ein Ankerelement mit `class="category career"` enthält.

```js
let actual = forumCategory(299);
assert.match(actual, /class=("|')category\s+career\1/);

// prevent hardcoding
actual = forumCategory(409);
assert.match(actual, /class=("|')category\s+feedback\1/);
```

`forumCategory(200)` sollte einen String zurückgeben, der ein Ankerelement mit dem Text `General` enthält.

```js
const actual = forumCategory(200);
assert.match(actual.trim(), /^<\s*a.+?>\s*General\s*<\/a>$/);
```

`forumCategory(200)` sollte einen String zurückgeben, der ein Ankerelement mit `href="https://forum.freecodecamp.org/c/general/200"` enthält.

```js
let actual = forumCategory(200);
assert.match(
  actual,
  /href=("|')https:\/\/forum\.freecodecamp\.org\/c\/general\/200/
);

actual = forumCategory(220);
assert.match(
  actual,
  /href=("|')https:\/\/forum\.freecodecamp\.org\/c\/general\/220/
);
```

`forumCategory(200)` sollte einen String zurückgeben, der ein Ankerelement mit `class="category general"` enthält.

```js
const actual = forumCategory(200);
assert.match(actual, /class=("|')category\s+general\1/);
```

Sie sollten eine Funktion namens `avatars` haben, die zwei Argumente annimmt.

```js
assert.isFunction(avatars);
assert.lengthOf(avatars, 2);
```

Die Funktion `avatars` sollte einen String zurückgeben, der durch das Verbinden von `img`-Elementen gebildet wird, eines für jeden Poster, der im Benutzerarray gefunden wurde.

```js
const posters = [{ user_id: 6 }, { user_id: 285941 }, { user_id: 170865 }];
const users = [
  {
    avatar_template:
      '/user_avatar/QuincyLarson_{size}.png',
    id: 6,
    name: 'Quincy Larson',
    username: 'QuincyLarson'
  },
  {
    avatar_template:
      '/user_avatar/jwilkins.oboe_{size}.png',
    id: 285941,
    name: 'Jessica Wilkins',
    username: 'jwilkins.oboe'
  },
  {
    avatar_template:
      '/user_avatar/ilenia_{size}.png',
    id: 170865,
    name: 'Ilenia',
    username: 'ilenia'
  },
  { id: 20 }
];
const actual = avatars(posters, users);
const matches = actual.match(/<\s*img\s+.+?>/g);
assert.lengthOf(matches, 3);
```

Jedes `img`-Element im String, der von der Funktion `avatars` zurückgegeben wird, sollte einen `alt`-Text mit dem Wert der Eigenschaft `name` des Posters haben.

```js
const posters = [{ user_id: 6 }, { user_id: 285941 }, { user_id: 170865 }];
const users = [
  {
    avatar_template:
      '/user_avatar/QuincyLarson_{size}.png',
    id: 6,
    name: 'Quincy Larson',
    username: 'QuincyLarson'
  },
  {
    avatar_template:
      '/user_avatar/jwilkins.oboe_{size}.png',
    id: 285941,
    name: 'Jessica Wilkins',
    username: 'jwilkins.oboe'
  },
  {
    avatar_template:
      '/user_avatar/ilenia_{size}.png',
    id: 170865,
    name: 'Ilenia',
    username: 'ilenia'
  },
  { id: 20 }
];
const actual = avatars(posters, users);
const matches = actual.match(/<\s*img\s+.+?>/g);

assert.match(matches[0], /alt=('|")Quincy Larson\1/);
assert.match(matches[1], /alt=('|")Jessica Wilkins\1/);
assert.match(matches[2], /alt=('|")Ilenia\1/);
```

Die Funktion `avatars` sollte die Größe jedes Avatars festlegen, indem sie auf die Eigenschaft `avatar_template` zugreift und `{size}` durch `30` ersetzt.

```js
const posters = [{ user_id: 6 }, { user_id: 285941 }, { user_id: 170865 }];
const users = [
  {
    avatar_template:
      '/user_avatar/QuincyLarson_{size}.png',
    id: 6,
    name: 'Quincy Larson',
    username: 'QuincyLarson'
  },
  {
    avatar_template:
      '/user_avatar/jwilkins.oboe_{size}.png',
    id: 285941,
    name: 'Jessica Wilkins',
    username: 'jwilkins.oboe'
  },
  {
    avatar_template:
      '/user_avatar/ilenia_{size}.png',
    id: 170865,
    name: 'Ilenia',
    username: 'ilenia'
  },
  { id: 20 }
];
const actual = avatars(posters, users);
assert.notMatch(actual, /\{size\}/);
assert.lengthOf(actual.match(/_30/g), 3);
```

Jedes `img`-Element sollte ein `src`-Attribut mit dem Wert `avatar_template` des passenden Benutzers haben. Wenn `avatar_template` mit `/` beginnt, fügen Sie `avatarUrl` direkt davor ein.

```js
const posters = [{ user_id: 6 }, { user_id: 285941 }, { user_id: 170865 }];
const users = [
  {
    avatar_template:
      '/user_avatar/QuincyLarson_{size}.png',
    id: 6,
    name: 'Quincy Larson',
    username: 'QuincyLarson'
  },
  {
    avatar_template:
      '/user_avatar/jwilkins.oboe_{size}.png',
    id: 285941,
    name: 'Jessica Wilkins',
    username: 'jwilkins.oboe'
  },
  {
    avatar_template:
      '/user_avatar/ilenia_{size}.png',
    id: 170865,
    name: 'Ilenia',
    username: 'ilenia'
  },
  { id: 20 }
];
const actual = avatars(posters, users);

const matches = actual.match(/<\s*img\s+.+?>/g);

assert.match(
  matches[0],
  /src=('|")https:\/\/cdn\.freecodecamp\.org\/curriculum\/forum-latest\/user_avatar\/QuincyLarson_30\.png\1/
);
assert.match(
  matches[1],
  /src=('|")https:\/\/cdn\.freecodecamp\.org\/curriculum\/forum-latest\/user_avatar\/jwilkins\.oboe_30\.png\1/
);
assert.match(
  matches[2],
  /src=('|")https:\/\/cdn\.freecodecamp\.org\/curriculum\/forum-latest\/user_avatar\/ilenia_30\.png\1/
);
```

Sie sollten eine Funktion namens `showLatestPosts` haben, die einen einzelnen Parameter annimmt.

```js
assert.isFunction(showLatestPosts);
assert.lengthOf(showLatestPosts, 1);
```

Sie sollten eine Funktion namens `fetchData` haben.

```js
assert.isFunction(fetchData);
```

Ihre Funktion `fetchData` sollte Daten von `https://cdn.freecodecamp.org/curriculum/forum-latest/latest.json` anfordern.

```js
const testArr = [];
const temp = fetch;
try {
  fetch = source => {
    testArr.push(source);
    return temp(source);
  };
  fetchData();
  assert.deepEqual(testArr, [
    'https://cdn.freecodecamp.org/curriculum/forum-latest/latest.json'
  ]);
} finally {
  fetch = temp;
}
```

Ihre Funktion `fetchData` sollte `showLatestPosts` aufrufen.

```js
const testArr = [];
const temp = showLatestPosts;
async () => {
  try {
    showLatestPosts = data => {
      testArr.push(data);
      return temp(data);
    };
    await fetchData();
    assert.isNotEmpty(testArr);
  } catch (err) {
    throw new Error(err);
  } finally {
    fetch = temp;
  }
};
```

Wenn ein Fehler auftritt, sollte Ihre Funktion `fetchData` den Fehler in der Konsole protokollieren, und zwar mit `console.log`.

```js
const testArr = [];
const temp1 = fetch;
const temp2 = console.log;
async () => {
  try {
    console.log = obj => {
      testArr.push(obj.toString());
    };
    fetch = source => {
      throw new Error('This is a test error');
    };
    await fetchData();
    assert.deepEqual(testArr, ['Error: This is a test error']);
  } finally {
    fetch = temp1;
    console.log = temp2;
  }
};
```

`showLatestPosts` sollte das innere HTML von `#posts-container` auf einen String setzen, der durch das Verbinden von `tr`-Elementen gebildet wird, eines für jeden Eintrag in `topics`.

```js
const data = {
  users: [
    {
      avatar_template:
        '/user_avatar/QuincyLarson_{size}.png',
      id: 6,
      name: 'Quincy Larson',
      username: 'QuincyLarson'
    },
    {
      avatar_template:
        '/user_avatar/jwilkins.oboe_{size}.png',
      id: 285941,
      name: 'Jessica Wilkins',
      username: 'jwilkins.oboe'
    },
    {
      avatar_template:
        '/user_avatar/ilenia_{size}.png',
      id: 170865,
      name: 'Ilenia',
      username: 'ilenia'
    }
  ],
  topic_list: {
    topics: [
      {
        bumped_at: '2024-04-15T16:01:26.403Z',
        category_id: 1,
        id: 684569,
        posters: [{ user_id: 6 }, { user_id: 170865 }, { user_id: 285941 }],
        posts_count: 8,
        slug: 'the-freecodecamp-podcast-is-back-now-with-video',
        title: 'The freeCodeCamp Podcast is back – now with video',
        views: 542
      },
      {
        bumped_at: '2024-04-19T13:52:03.523Z',
        category_id: 421,
        id: 686149,
        posters: [{ user_id: 170865 }],
        posts_count: 1,
        slug: 'problem-with-making-changes-to-styles-js',
        title: 'Problem with making changes to styles. (JS)',
        views: 9
      }
    ]
  }
};
const pContainer = document.getElementById('posts-container');
pContainer.innerHTML = '';
showLatestPosts(data);
assert.lengthOf(pContainer.querySelectorAll('tr'), 2);
```

Jedes `tr`-Element aus dem String, der von `showLatestPosts` zurückgegeben wird, sollte 5 `td`-Elemente enthalten.

```js
const data = {
  users: [
    {
      avatar_template:
        '/user_avatar/QuincyLarson_{size}.png',
      id: 6,
      name: 'Quincy Larson',
      username: 'QuincyLarson'
    },
    {
      avatar_template:
        '/user_avatar/jwilkins.oboe_{size}.png',
      id: 285941,
      name: 'Jessica Wilkins',
      username: 'jwilkins.oboe'
    },
    {
      avatar_template:
        '/user_avatar/ilenia_{size}.png',
      id: 170865,
      name: 'Ilenia',
      username: 'ilenia'
    }
  ],
  topic_list: {
    topics: [
      {
        bumped_at: '2024-04-15T16:01:26.403Z',
        category_id: 1,
        id: 684569,
        posters: [{ user_id: 6 }, { user_id: 170865 }, { user_id: 285941 }],
        posts_count: 8,
        slug: 'the-freecodecamp-podcast-is-back-now-with-video',
        title: 'The freeCodeCamp Podcast is back – now with video',
        views: 542
      },
      {
        bumped_at: '2024-04-19T13:52:03.523Z',
        category_id: 421,
        id: 686149,
        posters: [{ user_id: 170865 }],
        posts_count: 1,
        slug: 'problem-with-making-changes-to-styles-js',
        title: 'Problem with making changes to styles. (JS)',
        views: 9
      }
    ]
  }
};
const pContainer = document.getElementById('posts-container');
pContainer.innerHTML = '';
showLatestPosts(data);
assert.lengthOf(pContainer.querySelectorAll('tr:first-child>td'), 5);
assert.lengthOf(pContainer.querySelectorAll('tr:last-child>td'), 5);
```

Das erste `td`-Element jeder Tabellenzeile aus dem String, der von `showLatestPosts` zurückgegeben wird, sollte zwei Ankerelemente enthalten, das erste mit der Klasse `post-title`, einem `href` von `<forumTopicUrl><slug>/<id>`, einem Ankertext von `<title>` und das zweite, das durch Aufrufen von `forumCategory` mit `category_id` erhalten wird.

```js
const data = {
  users: [
    {
      avatar_template:
        '/user_avatar/QuincyLarson_{size}.png',
      id: 6,
      name: 'Quincy Larson',
      username: 'QuincyLarson'
    },
    {
      avatar_template:
        '/user_avatar/jwilkins.oboe_{size}.png',
      id: 285941,
      name: 'Jessica Wilkins',
      username: 'jwilkins.oboe'
    },
    {
      avatar_template:
        '/user_avatar/ilenia_{size}.png',
      id: 170865,
      name: 'Ilenia',
      username: 'ilenia'
    }
  ],
  topic_list: {
    topics: [
      {
        bumped_at: '2024-04-15T16:01:26.403Z',
        category_id: 1,
        id: 684569,
        posters: [{ user_id: 6 }, { user_id: 170865 }, { user_id: 285941 }],
        posts_count: 8,
        slug: 'the-freecodecamp-podcast-is-back-now-with-video',
        title: 'The freeCodeCamp Podcast is back – now with video',
        views: 542
      },
      {
        bumped_at: '2024-04-19T13:52:03.523Z',
        category_id: 421,
        id: 686149,
        posters: [{ user_id: 170865 }],
        posts_count: 1,
        slug: 'problem-with-making-changes-to-styles-js',
        title: 'Problem with making changes to styles. (JS)',
        views: 9
      }
    ]
  }
};
const pContainer = document.getElementById('posts-container');
pContainer.innerHTML = '';
showLatestPosts(data);
const anchors1 = pContainer.querySelectorAll('tr:first-child>td>a');
assert.lengthOf(anchors1, 2);

const anchors2 = pContainer.querySelectorAll('tr:last-child>td>a');
assert.lengthOf(anchors2, 2);

assert.equal(anchors1[0].classList[0], 'post-title');
assert.equal(
  anchors1[0].href,
  'https://forum.freecodecamp.org/t/the-freecodecamp-podcast-is-back-now-with-video/684569'
);
assert.equal(
  anchors1[0].innerText.trim(),
  'The freeCodeCamp Podcast is back – now with video'
);

assert.equal(anchors1[1].classList[0], 'category');
assert.equal(anchors1[1].classList[1], 'general');
assert.equal(anchors1[1].href, 'https://forum.freecodecamp.org/c/general/1');

assert.equal(anchors2[0].classList[0], 'post-title');
assert.equal(
  anchors2[0].href,
  'https://forum.freecodecamp.org/t/problem-with-making-changes-to-styles-js/686149'
);
assert.equal(
  anchors2[0].innerText.trim(),
  'Problem with making changes to styles. (JS)'
);

assert.equal(anchors2[1].classList[0], 'category');
assert.equal(anchors2[1].classList[1], 'javascript');
assert.equal(
  anchors2[1].href,
  'https://forum.freecodecamp.org/c/javascript/421'
);
```

Das zweite `td`-Element jeder Tabellenzeile aus dem String, der von `showLatestPosts` zurückgegeben wird, sollte die Bilder enthalten, die von der Funktion `avatars` zurückgegeben werden, die mit `posters` und `users` als Argumente aufgerufen wird, verschachtelt in einem `div`-Element mit der Klasse `avatar-container`.

```js
const data = {
  users: [
    {
      avatar_template:
        '/user_avatar/QuincyLarson_{size}.png',
      id: 6,
      name: 'Quincy Larson',
      username: 'QuincyLarson'
    },
    {
      avatar_template:
        '/user_avatar/jwilkins.oboe_{size}.png',
      id: 285941,
      name: 'Jessica Wilkins',
      username: 'jwilkins.oboe'
    },
    {
      avatar_template:
        '/user_avatar/ilenia_{size}.png',
      id: 170865,
      name: 'Ilenia',
      username: 'ilenia'
    }
  ],
  topic_list: {
    topics: [
      {
        bumped_at: '2024-04-15T16:01:26.403Z',
        category_id: 1,
        id: 684569,
        posters: [{ user_id: 6 }, { user_id: 170865 }, { user_id: 285941 }],
        posts_count: 8,
        slug: 'the-freecodecamp-podcast-is-back-now-with-video',
        title: 'The freeCodeCamp Podcast is back – now with video',
        views: 542
      },
      {
        bumped_at: '2024-04-19T13:52:03.523Z',
        category_id: 421,
        id: 686149,
        posters: [{ user_id: 170865 }],
        posts_count: 1,
        slug: 'problem-with-making-changes-to-styles-js',
        title: 'Problem with making changes to styles. (JS)',
        views: 9
      }
    ]
  }
};
const pContainer = document.getElementById('posts-container');
pContainer.innerHTML = '';
showLatestPosts(data);

const div1 = pContainer.querySelector('tr:first-child>td:nth-child(2)>div');
assert.equal(div1.classList[0], 'avatar-container');

const div2 = pContainer.querySelector('tr:last-child>td:nth-child(2)>div');
assert.equal(div2.classList[0], 'avatar-container');

const imgs1 = div1.querySelectorAll('img');
assert.lengthOf(imgs1, 3);
assert.equal(
  imgs1[0].src,
  'https://cdn.freecodecamp.org/curriculum/forum-latest/user_avatar/QuincyLarson_30.png'
);
assert.equal(imgs1[0].alt, 'Quincy Larson');
assert.equal(
  imgs1[1].src,
  'https://cdn.freecodecamp.org/curriculum/forum-latest/user_avatar/ilenia_30.png'
);
assert.equal(imgs1[1].alt, 'Ilenia');
assert.equal(
  imgs1[2].src,
  'https://cdn.freecodecamp.org/curriculum/forum-latest/user_avatar/jwilkins.oboe_30.png'
);
assert.equal(imgs1[2].alt, 'Jessica Wilkins');

const imgs2 = div2.querySelectorAll('img');
assert.lengthOf(imgs2, 1);
assert.equal(
  imgs2[0].src,
  'https://cdn.freecodecamp.org/curriculum/forum-latest/user_avatar/ilenia_30.png'
);
assert.equal(imgs2[0].alt, 'Ilenia');
```

Das dritte `td`-Element jeder Tabellenzeile aus dem String, der von `showLatestPosts` zurückgegeben wird, sollte die Anzahl der Antworten auf den Beitrag enthalten. _Hinweis:_ Verwenden Sie `posts_count - 1`.

```js
const data = {
  users: [
    {
      avatar_template:
        '/user_avatar/QuincyLarson_{size}.png',
      id: 6,
      name: 'Quincy Larson',
      username: 'QuincyLarson'
    },
    {
      avatar_template:
        '/user_avatar/jwilkins.oboe_{size}.png',
      id: 285941,
      name: 'Jessica Wilkins',
      username: 'jwilkins.oboe'
    },
    {
      avatar_template:
        '/user_avatar/ilenia_{size}.png',
      id: 170865,
      name: 'Ilenia',
      username: 'ilenia'
    }
  ],
  topic_list: {
    topics: [
      {
        bumped_at: '2024-04-15T16:01:26.403Z',
        category_id: 1,
        id: 684569,
        posters: [{ user_id: 6 }, { user_id: 170865 }, { user_id: 285941 }],
        posts_count: 8,
        slug: 'the-freecodecamp-podcast-is-back-now-with-video',
        title: 'The freeCodeCamp Podcast is back – now with video',
        views: 542
      },
      {
        bumped_at: '2024-04-19T13:52:03.523Z',
        category_id: 421,
        id: 686149,
        posters: [{ user_id: 170865 }],
        posts_count: 1,
        slug: 'problem-with-making-changes-to-styles-js',
        title: 'Problem with making changes to styles. (JS)',
        views: 9
      }
    ]
  }
};
const pContainer = document.getElementById('posts-container');
pContainer.innerHTML = '';
showLatestPosts(data);

assert.equal(
  pContainer.querySelector('tr:first-child>td:nth-child(3)').innerText,
  '7'
);

assert.equal(
  pContainer.querySelector('tr:last-child>td:nth-child(3)').innerText,
  '0'
);
```

Das vierte `td`-Element jeder Tabellenzeile aus dem String, der von `showLatestPosts` zurückgegeben wird, sollte die Anzahl der Ansichten des Beitrags enthalten.

```js
const data = {
  users: [
    {
      avatar_template:
        '/user_avatar/QuincyLarson_{size}.png',
      id: 6,
      name: 'Quincy Larson',
      username: 'QuincyLarson'
    },
    {
      avatar_template:
        '/user_avatar/jwilkins.oboe_{size}.png',
      id: 285941,
      name: 'Jessica Wilkins',
      username: 'jwilkins.oboe'
    },
    {
      avatar_template:
        '/user_avatar/ilenia_{size}.png',
      id: 170865,
      name: 'Ilenia',
      username: 'ilenia'
    }
  ],
  topic_list: {
    topics: [
      {
        bumped_at: '2024-04-15T16:01:26.403Z',
        category_id: 1,
        id: 684569,
        posters: [{ user_id: 6 }, { user_id: 170865 }, { user_id: 285941 }],
        posts_count: 8,
        slug: 'the-freecodecamp-podcast-is-back-now-with-video',
        title: 'The freeCodeCamp Podcast is back – now with video',
        views: 542
      },
      {
        bumped_at: '2024-04-19T13:52:03.523Z',
        category_id: 421,
        id: 686149,
        posters: [{ user_id: 170865 }],
        posts_count: 1,
        slug: 'problem-with-making-changes-to-styles-js',
        title: 'Problem with making changes to styles. (JS)',
        views: 9
      }
    ]
  }
};
const pContainer = document.getElementById('posts-container');
pContainer.innerHTML = '';
showLatestPosts(data);

assert.equal(
  pContainer.querySelector('tr:first-child>td:nth-child(4)').innerText,
  '542'
);

assert.equal(
  pContainer.querySelector('tr:last-child>td:nth-child(4)').innerText,
  '9'
);
```

Das fünfte `td`-Element jeder Tabellenzeile aus dem String, der von `showLatestPosts` zurückgegeben wird, sollte die seit der letzten Aktivität verstrichene Zeit enthalten, die mit der Funktion `timeAgo` generiert wurde.

```js
const data = {
  users: [
    {
      avatar_template:
        '/user_avatar/QuincyLarson_{size}.png',
      id: 6,
      name: 'Quincy Larson',
      username: 'QuincyLarson'
    },
    {
      avatar_template:
        '/user_avatar/jwilkins.oboe_{size}.png',
      id: 285941,
      name: 'Jessica Wilkins',
      username: 'jwilkins.oboe'
    },
    {
      avatar_template:
        '/user_avatar/ilenia_{size}.png',
      id: 170865,
      name: 'Ilenia',
      username: 'ilenia'
    }
  ],
  topic_list: {
    topics: [
      {
        bumped_at: '2024-04-15T16:01:26.403Z',
        category_id: 1,
        id: 684569,
        posters: [{ user_id: 6 }, { user_id: 170865 }, { user_id: 285941 }],
        posts_count: 8,
        slug: 'the-freecodecamp-podcast-is-back-now-with-video',
        title: 'The freeCodeCamp Podcast is back – now with video',
        views: 542
      },
      {
        bumped_at: '2024-04-19T13:52:03.523Z',
        category_id: 421,
        id: 686149,
        posters: [{ user_id: 170865 }],
        posts_count: 1,
        slug: 'problem-with-making-changes-to-styles-js',
        title: 'Problem with making changes to styles. (JS)',
        views: 9
      }
    ]
  }
};
const calcTime = time => {
  const currentTime = new Date();
  const lastPost = new Date(time);
  const timeDifference = currentTime - lastPost;
  const msPerMinute = 1000 * 60;
  const minutesAgo = Math.floor(timeDifference / msPerMinute);
  const hoursAgo = Math.floor(minutesAgo / 60);
  const daysAgo = Math.floor(hoursAgo / 24);
  return `${daysAgo}d ago`;
};
const pContainer = document.getElementById('posts-container');
pContainer.innerHTML = '';
showLatestPosts(data);

assert.equal(
  pContainer.querySelector('tr:first-child>td:nth-child(5)').innerText,
  calcTime('2024-04-15T16:01:26.403Z')
);

assert.equal(
  pContainer.querySelector('tr:last-child>td:nth-child(5)').innerText,
  calcTime('2024-04-19T13:52:03.523Z')
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>fCC Forum Leaderboard</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <nav>
        <img
          class="fcc-logo"
          src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
          alt="freeCodeCamp logo"
        />
      </nav>
      <h1 class="title">Latest Topics</h1>
    </header>
    <main>
      <div class="table-wrapper">
        <table>
          <thead>
            <tr>
              <th id="topics">Topics</th>
              <th id="avatars">Avatars</th>
              <th id="replies">Replies</th>
              <th id="views">Views</th>
              <th id="activity">Activity</th>
            </tr>
          </thead>
          <tbody id="posts-container"></tbody>
        </table>
      </div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

.fcc-logo {
  width: 210px;
  display: block;
  margin: auto;
}

.title {
  margin: 25px 0;
  text-align: center;
  color: var(--light-grey);
}

.table-wrapper {
  padding: 0 25px;
  overflow-x: auto;
}

table {
  width: 100%;
  color: var(--dark-grey);
  margin: auto;
  table-layout: fixed;
  border-collapse: collapse;
  overflow-x: scroll;
}

#topics {
  text-align: start;
  width: 60%;
}

th {
  border-bottom: 2px solid var(--dark-grey);
  padding-bottom: 10px;
  font-size: 1.3rem;
}

td:not(:first-child) {
  text-align: center;
}

td {
  border-bottom: 1px solid var(--dark-grey);
  padding: 20px 0;
}

.post-title {
  font-size: 1.2rem;
  color: var(--medium-grey);
  text-decoration: none;
}

.category {
  padding: 3px;
  color: var(--black);
  text-decoration: none;
  display: block;
  width: fit-content;
  margin: 10px 0 10px;
}

.career {
  background-color: var(--salmon-color);
}

.feedback,
.html-css {
  background-color: var(--light-blue);
}

.support {
  background-color: var(--light-orange);
}

.general {
  background-color: var(--light-green);
}

.javascript {
  background-color: var(--golden-yellow);
}

.backend {
  background-color: var(--gold);
}

.python {
  background-color: var(--green);
}

.motivation {
  background-color: var(--peach);
}

.avatar-container {
  display: flex;
  justify-content: center;
  gap: 10px;
  flex-wrap: wrap;
}

.avatar-container img {
  width: 30px;
  height: 30px;
}

@media (max-width: 750px) {
  .table-wrapper {
    padding: 0 15px;
  }

  table {
    width: 700px;
  }

  th {
    font-size: 1.2rem;
  }

  .post-title {
    font-size: 1.1rem;
  }
}
```

```js
const forumLatest =
  'https://cdn.freecodecamp.org/curriculum/forum-latest/latest.json';
const forumTopicUrl = 'https://forum.freecodecamp.org/t/';
const forumCategoryUrl = 'https://forum.freecodecamp.org/c/';
const avatarUrl = 'https://cdn.freecodecamp.org/curriculum/forum-latest';

const allCategories = {
  299: { category: 'Career Advice', className: 'career' },
  409: { category: 'Project Feedback', className: 'feedback' },
  417: { category: 'freeCodeCamp Support', className: 'support' },
  421: { category: 'JavaScript', className: 'javascript' },
  423: { category: 'HTML - CSS', className: 'html-css' },
  424: { category: 'Python', className: 'python' },
  432: { category: 'You Can Do This!', className: 'motivation' },
  560: { category: 'Back-End Development', className: 'backend' }
};
```

# --solutions--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>fCC Forum Leaderboard</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <nav>
        <img
          class="fcc-logo"
          src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
          alt="freeCodeCamp logo"
        />
      </nav>
      <h1 class="title">Latest Topics</h1>
    </header>
    <main>
      <div class="table-wrapper">
        <table>
          <thead>
            <tr>
              <th id="topics">Topics</th>
              <th id="avatars">Avatars</th>
              <th id="replies">Replies</th>
              <th id="views">Views</th>
              <th id="activity">Activity</th>
            </tr>
          </thead>
          <tbody id="posts-container"></tbody>
        </table>
      </div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

.fcc-logo {
  width: 210px;
  display: block;
  margin: auto;
}

.title {
  margin: 25px 0;
  text-align: center;
  color: var(--light-grey);
}

.table-wrapper {
  padding: 0 25px;
  overflow-x: auto;
}

table {
  width: 100%;
  color: var(--dark-grey);
  margin: auto;
  table-layout: fixed;
  border-collapse: collapse;
  overflow-x: scroll;
}

#topics {
  text-align: start;
  width: 60%;
}

th {
  border-bottom: 2px solid var(--dark-grey);
  padding-bottom: 10px;
  font-size: 1.3rem;
}

td:not(:first-child) {
  text-align: center;
}

td {
  border-bottom: 1px solid var(--dark-grey);
  padding: 20px 0;
}

.post-title {
  font-size: 1.2rem;
  color: var(--medium-grey);
  text-decoration: none;
}

.category {
  padding: 3px;
  color: var(--black);
  text-decoration: none;
  display: block;
  width: fit-content;
  margin: 10px 0 10px;
}

.career {
  background-color: var(--salmon-color);
}

.feedback,
.html-css {
  background-color: var(--light-blue);
}

.support {
  background-color: var(--light-orange);
}

.general {
  background-color: var(--light-green);
}

.javascript {
  background-color: var(--golden-yellow);
}

.backend {
  background-color: var(--gold);
}

.python {
  background-color: var(--green);
}

.motivation {
  background-color: var(--peach);
}

.avatar-container {
  display: flex;
  justify-content: center;
  gap: 10px;
  flex-wrap: wrap;
}

.avatar-container img {
  width: 30px;
  height: 30px;
}

@media (max-width: 750px) {
  .table-wrapper {
    padding: 0 15px;
  }

  table {
    width: 700px;
  }

  th {
    font-size: 1.2rem;
  }

  .post-title {
    font-size: 1.1rem;
  }
}
```

```js
const forumLatest =
  'https://cdn.freecodecamp.org/curriculum/forum-latest/latest.json';
const forumTopicUrl = 'https://forum.freecodecamp.org/t/';
const forumCategoryUrl = 'https://forum.freecodecamp.org/c/';
const avatarUrl = 'https://cdn.freecodecamp.org/curriculum/forum-latest';

const postsContainer = document.getElementById('posts-container');

const allCategories = {
  299: { category: 'Career Advice', className: 'career' },
  409: { category: 'Project Feedback', className: 'feedback' },
  417: { category: 'freeCodeCamp Support', className: 'support' },
  421: { category: 'JavaScript', className: 'javascript' },
  423: { category: 'HTML - CSS', className: 'html-css' },
  424: { category: 'Python', className: 'python' },
  432: { category: 'You Can Do This!', className: 'motivation' },
  560: { category: 'Back-End Development', className: 'backend' }
};

const forumCategory = id => {
  let selectedCategory = {};

  if (allCategories.hasOwnProperty(id)) {
    const { className, category } = allCategories[id];

    selectedCategory.className = className;
    selectedCategory.category = category;
  } else {
    selectedCategory.className = 'general';
    selectedCategory.category = 'General';
    selectedCategory.id = 1;
  }
  const url = `${forumCategoryUrl}${selectedCategory.className}/${id}`;
  const linkText = selectedCategory.category;
  const linkClass = `category ${selectedCategory.className}`;

  return `<a href="${url}" class="${linkClass}" target="_blank">
    ${linkText}
  </a>`;
};

const timeAgo = time => {
  const currentTime = new Date();
  const lastPost = new Date(time);

  const timeDifference = currentTime - lastPost;
  const msPerMinute = 1000 * 60;

  const minutesAgo = Math.floor(timeDifference / msPerMinute);
  const hoursAgo = Math.floor(minutesAgo / 60);
  const daysAgo = Math.floor(hoursAgo / 24);

  if (minutesAgo < 60) {
    return `${minutesAgo}m ago`;
  }

  if (hoursAgo < 24) {
    return `${hoursAgo}h ago`;
  }

  return `${daysAgo}d ago`;
};

const viewCount = views => {
  const thousands = Math.floor(views / 1000);

  if (views >= 1000) {
    return `${thousands}k`;
  }

  return views;
};

const avatars = (posters, users) => {
  return posters
    .map(poster => {
      const user = users.find(user => user.id === poster.user_id);
      if (user) {
        const avatar = user.avatar_template.replace(/{size}/, 30);
        const userAvatarUrl = avatar.startsWith('/user_avatar/')
          ? avatarUrl.concat(avatar)
          : avatar;
        return `<img src="${userAvatarUrl}" alt="${user.name}" />`;
      }
    })
    .join('');
};

const fetchData = async () => {
  try {
    const res = await fetch(forumLatest);
    const data = await res.json();
    showLatestPosts(data);
  } catch (err) {
    console.log(err);
  }
};

fetchData();

const showLatestPosts = data => {
  const { topic_list, users } = data;
  const { topics } = topic_list;

  postsContainer.innerHTML = topics
    .map(item => {
      const {
        id,
        title,
        views,
        posts_count,
        slug,
        posters,
        category_id,
        bumped_at
      } = item;

      return `
    <tr>
      <td>
       <a class="post-title" target="_blank" href="${forumTopicUrl}${slug}/${id}">
        ${title}
       </a>
        ${forumCategory(category_id)}
      </td>
      <td>
        <div class="avatar-container">
          ${avatars(posters, users)}
        </div>
      </td>
      <td>${posts_count - 1}</td>
      <td>${viewCount(views)}</td>
      <td>${timeAgo(bumped_at)}</td>
    </tr>`;
    })
    .join('');
};
```
