---
id: 60a3e3396c7b40068ad6996f
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

`body` 内に `div` 要素を追加してください。

`class` 属性を `canvas` に設定してください。

This will act as the canvas for your painting.

# --hints--

Your code should have a `div` element.

```js
assert(document.querySelectorAll('div').length === 1)
```

The `div` element should have a `class` with the value `canvas`.

```js
assert(document.querySelector('div').className.split(' ').includes('canvas'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
