---
id: 5dc17d3bf86c76b9248c6eb4
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Елемент `p` використовують, щоб створити абзац тексту на вебсайті. Створіть елемент `p` під елементом `h2` та надайте йому такий текст:

`Everyone loves cute cats online!`

# --hints--

Елемент `p` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('p'));
```

Елемент `p` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.match(code, /<\/p\>/);
```

Текстом елемента `p` має бути `Everyone loves cute cats online!` Ви або не написали текст, або маєте друкарську помилку.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  ?.innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

Елемент `p` має бути під елементом `h2`. Вони знаходяться в неправильному порядку.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H2'), collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>

--fcc-editable-region--
  </body>
</html>
```
