---
id: 5dc174fcf86c76b9248c6eb2
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті зі створення застосунку із фотографіями котів ви дізнаєтесь, як працювати з базовими HTML-елементами, серед яких заголовки, абзаци та списки.

Перший елемент, про який ви дізнаєтеся, — елемент `h1`. Елемент `h1` є елементом, який використовується для головного заголовка вебсторінки.

```html
<h1>This is a main heading</h1>
```

Додайте елемент `h1` з текстом `CatPhotoApp` і зверніть увагу на зміну в попередньому перегляді.

Коли закінчите, натисніть на кнопку «Перевірити код», щоб переконатись, що все правильно.

# --hints--

Код повинен містити текст `CatPhotoApp`. Можливо, потрібно перевірити написання.

```js
assert.match(code, /catphotoapp/i);
```

Елемент `h1` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Елемент `h1` повинен мати кінцевий тег. Кінцеві теги мають такий синтаксис: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Текстом елемента `h1` має бути `CatPhotoApp`. Ви маєте друкарську помилку, пропустили текст або він відсутній між тегами елемента `h1`.

```js
assert.equal(document.querySelector('h1')?.innerText.toLowerCase(), 'catphotoapp');
```

Здається, ви використовуєте розширення браузера, яке змінює сторінку. Переконайтеся, що вимкнули усі розширення браузера.

```js
if(__checkForBrowserExtensions){
  assert.isAtMost(document.querySelectorAll('script').length, 2);
  assert.equal(document.querySelectorAll('style').length, 1);
  assert.equal(document.querySelectorAll('link').length, 0);
}
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
