---
id: 5dc174fcf86c76b9248c6eb2
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті зі створення застосунку із фотографіями котів ви продовжите працювати з базовими HTML-елементами, серед яких заголовки, абзаци та списки.

Для початку додайте елемент `h1` з текстом `CatPhotoApp`.

# --hints--

Код повинен містити текст `CatPhotoApp`. Можливо, потрібно перевірити написання.

```js
assert.match(code, /catphotoapp/i);
```

Елемент `h1` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Елемент `h1` повинен мати кінцевий тег. Кінцеві теги мають такий синтаксис: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Текстом елемента `h1` має бути `CatPhotoApp`. Ви маєте друкарську помилку, пропустили текст або він відсутній між тегами елемента `h1`.

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
