---
id: 5dc1798ff86c76b9248c6eb3
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Під елементом `h1` додайте елемент `h2` з таким текстом:

`Cat Photos`

# --hints--

Елемент `h1` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Елемент `h1` повинен мати кінцевий тег. Кінцеві теги мають такий синтаксис: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Ви повинні мати тільки один елемент `h1`. Видаліть все зайве.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Текстом елемента `h1` має бути `'CatPhotoApp'`. Ви або не написали текст, або маєте друкарську помилку.

```js
assert.equal(document.querySelector('h1')?.innerText.toLowerCase(), 'catphotoapp');
```

Елемент `h2` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('h2'));
```

Елемент `h2` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.match(code, /<\/h2\>/);
```

Текстом елемента `h2` має бути `Cat Photos`. Між початковими та кінцевими тегами елемента `h2` розмістіть лише текст `Cat Photos`.

```js
assert.equal(document.querySelector('h2')?.innerText.toLowerCase(), 'cat photos');
```

Елемент `h2` має бути під елементом `h1`. Елемент `h1` має більшу важливість, тому має бути над елементом `h2`.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H1'), collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>

--fcc-editable-region--
  </body>
</html>
```
