---
id: 5dc17dc8f86c76b9248c6eb5
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Коментарі дозволяють залишати повідомлення без впливу на відтворення у браузері. Вони також дозволяють зробити код неактивним. Коментар в HTML починається із `<!--`, містить будь-яку кількість рядків тексту та закінчується на `-->`.

Ось приклад коментаря з текстом `TODO: Remove h1`:

```html
<!-- TODO: Remove h1 -->
```

Додайте коментар над елементом `p` з таким текстом:

`TODO: Add link to cat photos`

# --hints--

Коментар має починатися з `<!--`. Не вистачає одного чи більше символів, що визначають початок коментаря.

```js
assert.match(code, /<!--/);
```

Коментар має закінчуватися на `-->`.  Не вистачає одного чи більше символів, що визначають кінець коментаря.

```js
assert.match(code, /-->/);
```

Код не повинен містити зайвих початкових/кінцевих символів коментаря. Ви маєте зайві символи `<!--` або `-->`, що відтворюються в браузері.

```js
const noSpaces = code.replace(/\s/g, '');
assert.isBelow(noSpaces.match(/<!--/g)?.length, 2)
assert.isBelow(noSpaces.match(/-->/g)?.length, 2);
```

Коментар має бути над елементом `p`. Вони знаходяться в неправильному порядку.

```js
assert.match(
  code.replace(/\s/g, ''),
  /<!--(.*?)--><p>everyonelovescutecatsonline!<\/p>/i
);
```

Коментар має містити текст `TODO: Add link to cat photos`.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--

    <p>Everyone loves cute cats online!</p>

--fcc-editable-region--
  </body>
</html>
```
