---
id: 5dc23f9bf86c76b9248c6eba
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

За допомогою елемента `img` можна додати зображення на вебсайт. Елементи `img` мають початковий тег, але не мають кінцевого. Елемент без кінцевого тега відомий як <dfn>порожній елемент</dfn>.

Додайте елемент `img` під елементом `p`. Наразі жодне зображення не буде відтворюватись в браузері.

# --hints--

Елемент `img` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('img'));
```

Елемент `img` не повинен мати кінцевого тегу `</img>`.

```js
assert.notMatch(code, /<\/img\>/);
```

Ви повинні мати тільки один елемент `img`. Видаліть все зайве.

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

Елемент `img` має бути під елементом `p`. Вони знаходяться в неправильному порядку.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```
