---
id: 5dc23991f86c76b9248c6eb8
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

У попередньому кроці ви розмістили елементи `h1`, `h2`, коментар та `p` в межах елемента `main`. Це називається *вкладеністю*. Вкладені елементи потрібно розміщувати на два пробіли праворуч від елемента, в якому вони вкладені. Цей інтервал називається відступом; його використовують, щоб полегшити читання HTML.

Ось приклад вкладеності та відступу:

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Відступ елементів `h1`, `h2` та коментаря на два пробіли більший, ніж елемента `main` під ними. Використайте пробіл на клавіатурі, щоб додати ще два пробіли перед елементом `p`, щоб він також мав належний відступ.

# --hints--

Код повинен мати елемент `h2` з текстом `Cat Photos`. Можливо, ви випадково видалили його, в ньому відсутні початковий/кінцевий теги або змінений текст.

```js
assert.exists(document.querySelector('h2'));
assert.match(code, /<\/h2\>/);
assert.equal(document.querySelector('h2')?.innerText.toLowerCase(),'cat photos')
```

Не додавайте елементи `ul` або `li` з прикладу.

```js
assert.notExists(document.querySelector('ul'));
assert.notExists(document.querySelector('li'));
```

Не змінюйте відступ в рядку з елементом `h2`. Його початковий тег має починатися через 6 пробілів від початку рядка. Ви можете перезапустити крок, щоб відновити початковий відступ.

```js
assert.match(code.toLowerCase(), /<\/h1\>\s*\n\s{6}<h2>/);
```

Код повинен мати коментар. Ви видалили коментар з попереднього кроку.

```js
assert.match(code, /<!--.*-->/);
```

Текстом коментаря має бути `TODO: Add link to cat photos`. Не змінюйте текст та пробіли коментаря.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i);
```

Не змінюйте відступ в рядку з коментарем. Його початковий тег має починатися через 6 пробілів від початку рядка. Ви можете перезапустити крок, щоб відновити початковий відступ.

```js
assert.match(
  code
    .toLowerCase(),
    /<\/h2>\s*\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/
);
```

Код повинен мати елемент `p`. Ви видалили елемент `p` з попереднього кроку.

```js
assert.exists(document.querySelector('p'));
```

Текстом елемента `p` має бути `Everyone loves cute cats online!` Не змінюйте текст, відступи чи пунктуацію елемента `p`.

```js
assert.match(
  document
    .querySelector('p')
    ?.innerText.toLowerCase(),
    /everyone\s+loves\s+cute\s+cats\s+online!$/
);
```

Початковий тег елемента `p` повинен мати відступ, який відповідає елементу `h2` та коментарю. Його початковий тег має починатися через 6 пробілів від початку рядка.

```js
assert.match(code.toLowerCase(), /-->\s*\n\s{6}<p>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    --fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
