---
id: 5dfb6a35eacea3f48c6300b4
title: Крок 26
challengeType: 0
dashedName: step-26
---

# --description--

Підпис зображення (`figcaption`) використовують, щоб описати зображення, яке міститься в елементі `figure`.

Ось приклад елемента `figcaption` з підписом `A cute cat`:

```html
<figure>
  <img src="image.jpg" alt="A description of the image">
  <figcaption>A cute cat</figcaption>
</figure>
```

Після зображення, вкладеного в елементі `figure`, додайте елемент `figcaption` з таким текстом:

`Cats love lasagna.`

# --hints--

Елемент `figcaption` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('figcaption'));
```

Елемент `figcaption` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.match(code, /<\/figcaption\>/);
```

Елемент `figcaption` має бути вкладеним в елементі `figure`.

```js
assert.exists(document.querySelector('figure > figcaption'));
```

Елемент `img` «lasagna» має бути вкладеним в елементі `figure`.

```js
assert.exists(document.querySelector('figure > img'))
assert.equal(
    document.querySelector('figure > img')?.getAttribute('src').toLowerCase(),
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

Елемент `figcaption`, вкладений в елемент `figure`, має бути під елементом `img`. Вони знаходяться в неправильному порядку.

```js
assert.equal(
  document.querySelector('figcaption')?.previousElementSibling.nodeName, 'IMG'
);
```

Текстом елемента `figcaption` має бути `Cats love lasagna.` Ви або не написали текст, або маєте друкарську помилку.

```js
assert.match(
  document.querySelector('figcaption')?.innerText?.trim(), /Cats love lasagna.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
