---
id: 5ef9b03c81a63668521804d2
title: Крок 29
challengeType: 0
dashedName: step-29
---

# --description--

Код для впорядкованого списку (`ol`) подібний до невпорядкованого списку, але елементи списку в упорядкованому списку відтворюються разом з числами.

Під елементом `h3` додайте впорядкований список з такими трьома елементами:

`flea treatment` `thunder` `other cats`

# --hints--

Елемент `ol` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('ol'));
```

Елемент `ol` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.match(code, /<\/ol>/);
```

Елемент `ol` має бути над кінцевим тегом другого елемента `section`. Вони знаходяться в неправильному порядку.

```js
assert.equal(document.querySelectorAll('main > section')[1]?.lastElementChild.nodeName, 'OL');
```

Три елементи `li` мають бути вкладеними в елементі `ol`.

```js
assert.lengthOf(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ), 3
);
```

У вас мають бути три елементи `li` з текстом `flea treatment`, `thunder` та `other cats` у будь-якому порядку.

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText.toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

Ви повинні мати тільки один елемент `ol`.

```js 
assert.lengthOf([...document.querySelectorAll('ol')], 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--
        <h3>Top 3 things cats hate:</h3>

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```
