---
id: 66731cd027ef3acb155669f5
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will learn how to work with basic HTML elements like headings, paragraphs, and lists by building a cat photo app.

Перший елемент, про який ви дізнаєтеся, — елемент `h1`. Елемент `h1` є елементом заголовка, який використовується для головного заголовка вебсторінки.

```html
<h1>This is a main heading</h1>
```

Add an `h1` element with the text of `CatPhotoApp` and watch the change in the browser preview.

Коли закінчите, натисніть на кнопку «Перевірити код», щоб переконатись, що все правильно.

# --hints--

Текст `CatPhotoApp` повинен бути в коді. Можливо, доведеться перевірити написання.

```js
assert.match(code, /catphotoapp/i);
```

Елемент `h1` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Елемент `h1` повинен мати кінцевий тег. Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Текстом елемента `h1` має бути `CatPhotoApp`. Ви пропустили текст, маєте друкарську помилку, або його немає між початковим та кінцевим тегами елемента `h1`.

```js
assert.equal(document.querySelector('h1').innerText.toLowerCase(), 'catphotoapp');
```

Здається, ви використовуєте розширення браузера, яке змінює сторінку. Переконайтеся, що вимкнули усі розширення браузера.

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
