---
id: 6690b9626a66f51257ae5a91
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

The `h1` through `h6` heading elements are used to signify the importance of content below them. The lower the number, the higher the importance, so `h2` elements have less importance than `h1` elements.

```html
<h1>most important heading element</h1>
<h2>second most important heading element</h2>
<h3>third most important heading element</h3>
<h4>fourth most important heading element</h4>
<h5>fifth most important heading element</h5>
<h6>least important heading element</h6>
```

Only use one `h1` element per page and place lower-importance headings below higher-importance headings.

Below the `h1` element, add an `h2` element with this text:

`Cat Photos`

# --hints--

Елемент `h1` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Елемент `h1` повинен мати кінцевий тег. Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

You should only have one `h1` element. Remove the extra.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Текстом елемента `h1` має бути `'CatPhotoApp'`. Ви або не написали текст, або маєте друкарську помилку.

```js
assert.equal(document.querySelector('h1').innerText.toLowerCase(), 'catphotoapp');
```

Елемент `h2` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('h2'));
```

Елемент `h2` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.match(code, /<\/h2\>/);
```

Your `h2` element's text should be `Cat Photos`. Only place the text `Cat Photos` between the opening and closing `h2` tags.

```js
assert.equal(document.querySelector('h2').innerText.toLowerCase(), 'cat photos');
```

Your `h2` element should be below the `h1` element. The `h1` element has greater importance and must be above the `h2` element.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H1'), collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>

--fcc-editable-region--
  </body>
</html>
```
