---
id: 6690beec809b4f184cab6914
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Ви можете скерувати користувача на іншу сторінку за допомогою елемента посилання (`a`).

Ось приклад посилання на `https://www.freecodecamp.org`:

```html
<a href="https://www.freecodecamp.org"></a>
```

Додайте елемент прив’язки після абзацу, який посилає на `https://freecatphotoapp.com`. На цей момент посилання не з’явиться в попередньому перегляді.

# --hints--

Елемент посилання (`a`) повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('a'));
```

Елемент посилання (`a`) повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.match(code, /<\/a\>/);
```

Елемент посилання (`a`) повинен бути під елементом `p`. Вони знаходяться в неправильному порядку.

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('A'));
```

Елемент посилання (`a`) не має атрибута `href`. Перевірте, чи є пробіл після назви початкового тега та/або є пробіли перед усіма назвами атрибутів.

```js
assert.isTrue(document.querySelector('a').hasAttribute('href'));
```

Елемент прив’язки (`a`) повинен посилати на `https://freecatphotoapp.com`. Ви або не написали URL-адресу, або маєте друкарську помилку.

```js
assert.equal(
  document.querySelector('a').getAttribute('href'),
    'https://freecatphotoapp.com'
);
```

Хоча ви встановили атрибут `href` елемента прив’язки ('a') як правильне посилання, краще завжди писати значення атрибута в лапках.

```js
assert.notMatch(
  code, /\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```


