---
id: 6690bd5ad96139157eedf697
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

In the previous step, you put the `h1`, `h2`, comment, and `p` elements inside the `main` element. Це називається *вкладеність*. Вкладені елементи потрібно розміщувати на два пробіли праворуч від елемента, в який вони вкладені. Цей інтервал називається відступом та використовується для полегшення читання HTML.

Ось приклад вкладеності та відступу:

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Елемент `h1`, елемент `h2` та коментар мають відступ на два пробіли більше, ніж елемент `main` в коді нижче. Скористайтеся пробілом на клавіатурі, щоб додати ще два пробіли перед елементом `p`, щоб він також мав належний відступ.

# --hints--

Код повинен мати елемент `h2` з текстом `Cat Photos`. Можливо, ви випадково видалили його, в ньому відсутні обидва початковий та кінцевий теґи, або змінився текст.

```js
assert.exists(document.querySelector('h2'));
assert.match(code, /<\/h2\>/);
assert.equal(document.querySelector('h2').innerText.toLowerCase(),'cat photos')
```

Ви не повинні додавати елементи `ul` або `li` з прикладу.

```js
assert.notExists(document.querySelector('ul'));
assert.notExists(document.querySelector('li'));
```

Не змінюйте відступ в рядку з елементом `h2`. Його початковий тег має починатися через 6 пробілів від початку рядка. Ви можете перезапустити крок, щоб відновити оригінальний відступ.

```js
assert.match(code.toLowerCase(), /<\/h1\>\s*\n\s{6}<h2>/);
```

Код повинен мати коментар. Ви видалили коментар з попереднього кроку.

```js
assert.match(code, /<!--.*-->/);
```

Текст коментаря повинен бути `TODO: Add link to cat photos`. Не змінюйте текст та пробіли коментаря.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i);
```

Не змінюйте відступ в рядку з елементом-коментарем. Його початковий тег має починатися через 6 пробілів від початку рядка. Ви можете перезапустити крок, щоб відновити оригінальний відступ.

```js
assert.match(
  code
    .toLowerCase(),
    /<\/h2>\s*\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/
);
```

Код повинен мати елемент `p`. Ви видалили елемент `p` з попереднього кроку.

```js
assert.exists(document.querySelector('p'));
```

Текст елемента `p` має бути `See more cat photos in our gallery.` Не змінюйте текст, відступи чи пунктуацію елемента `p`.

```js
assert.match(
  document
    .querySelector('p')
    .innerText.toLowerCase(),
    /see\s+more\s+cat\s+photos\s+in\s+our\s+gallery\.?$/
);
```

Початковий тег `p` повинен мати відступ, який відповідає елементам `h2` та коментарю. Його початковий тег має починатися через 6 пробілів від початку рядка.

```js
assert.match(code.toLowerCase(), /-->\s*\n\s{6}<p>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    --fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
    <p>See more cat photos in our gallery.</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```


