---
id: 66984eef66cc1752d93b53f5
title: Крок 36
challengeType: 0
dashedName: step-36
---

# --description--

Зверніть увагу, що все, що ви додали до сторінки, знаходиться всередині елемента `body`. Всі елементи вмісту сторінки, які мають відтворюватися на сторінці, йдуть в елемент `body`. Однак інша важлива інформація йде всередину елемента `head`.

Елемент `head` використовується для зберігання метаданих про документ (серед яких його назва, посилання на таблиці стилів та скрипти). Метадані — це інформація про сторінку, яка не відображається безпосередньо на сторінці.

Додайте елемент `head` над елементом `body`.

# --hints--

Ви або видалили елемент `body`, або в ньому відсутній початковий чи кінцевий тег.

```js
assert.exists(document.querySelector('body'));
assert.match(code, /<\/body>/);
```

Елемент `head` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.match(code, /\<head\>/);
```

Елемент `head` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.match(code, /\<\/head\>/);
```

Елемент `head` повинен бути над початковим тегом елемента `body`. Ви поклали його в іншому місці.

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<\/head\>\<body\>/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<html>

  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```


