---
id: 66984e8ab54bee517207be42
title: Крок 34
challengeType: 0
dashedName: step-34
---

# --description--

Вкладіть елемент `p` із текстом `No Copyright - freeCodeCamp.org` в елемент `footer`.

# --hints--

Ви або видалили елемент `footer`, або в ньому відсутній початковий чи кінцевий тег.

```js
assert.exists(document.querySelector('footer'));
assert.match(code, /<\/footer>/);
```

Ваш елемент `footer` повинен мати елемент `p`. Переконайтеся, що додали початковий та кінцевий теги для елемента `p`.

```js
assert.exists(document.querySelector('footer > p'));
```

Елемент `footer` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
const pElemClosingTags = code.match(/<\/p\>/g);
assert.exists(pElemClosingTags);
assert.lengthOf(pElemClosingTags, 2);
```

Текстом елемента `p` має бути `No Copyright - freeCodeCamp.org`. Ви пропустили текст, маєте друкарську помилку, або його немає між початковим та кінцевим теґами елемента `p`.

```js
const extraSpacesRemoved = document.querySelectorAll('footer > p')[0].innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /No Copyright - freeCodeCamp\.org$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
--fcc-editable-region--
    <footer>

    </footer>
--fcc-editable-region--
  </body>
</html>
```


