---
id: 66984218720e29393b1adad2
title: Крок 16
challengeType: 0
dashedName: step-16
---

# --description--

Перш ніж додавати будь-який новий вміст, вам слід використати елемент `section`, щоб відокремити вміст фотографій котів від майбутнього вмісту.

Елемент `section` використовують, щоб визначити розділи в документі (наприклад, частини, верхні та нижні колонтитули або будь-який інший розділ документа). Це семантичний елемент, який допомагає з SEO та доступністю.

```html
<section>
  <h2>Section Title</h2>
  <p>Section content...</p>
</section>
```

Take your `h2`, comment, `p`, and anchor (`a`) elements and nest them in a `section` element.

# --hints--

Елемент `section` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('section'));
```

Елемент `section` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.match(code, /<\/section\s*>/i);
```

Весь елемент `section` повинен бути між початковим та кінцевим тегами елемента `main`.

```js
assert.equal(document.querySelector('section').parentNode.nodeName, 'MAIN');
```

The existing `h2`, comment, `p` element, and anchor (`a`) element should be between the opening and closing tags of the `section` element.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert.lengthOf(foundElements, 3)
```

Елемент `h1` не повинен бути вкладеним в елементі `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    --fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>

      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

    </main>
--fcc-editable-region--
  </body>
</html>
```


