---
id: 6698429ba5761b39f8aad814
title: Step 17
challengeType: 0
dashedName: step-17
---

# --description--

It is time to add a new section. Додайте другий елемент `section` під наявним елементом `section`.

# --hints--

Елемент `section` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.isAtLeast(document.querySelectorAll('section').length, 2);
```

You should only add one opening `section` tag. Please remove any extras.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
```

Елемент `section` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.isAtLeast(code.match(/<\/section>/g).length, 2);
```

You should only add one closing `section` tag. Please remove any extras.

```js
assert.lengthOf(code.match(/<\/section>/g), 2);
```

The second `section` element should not be nested in the first `section` element.

```js
const childrenOf1stSection = [
  ...document.querySelector('main > section').children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.isEmpty(foundElems);
```

Both `section` elements should be between the opening and closing tags of the `main` element.

```js
const childrenOfMain = [...document.querySelector('main').children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.lengthOf(foundElems, 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>

--fcc-editable-region--
    </main>
  </body>
</html>
```


