---
id: 669842cee55b683aa13ee758
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

Within the second `section` element, add a new `h2` element with the text `Cat Lists`.

# --hints--

Елемент `section` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
assert.lengthOf(code.match(/<\/section>/g), 2);
```

Елемент `h2` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('h2'), 2);
```

Елемент `h2` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.lengthOf(code.match(/<\/h2\>/g), 2);
```

Your second `h2` element should be right above the second `section` element's closing tag. It is not in the correct position.

```js
const secondSection = document.querySelectorAll('section')[1];
assert.equal(secondSection.lastElementChild.nodeName, 'H2');
```

The second `h2` element should have the text `Cat Lists`. Ви або не написали текст, або маєте друкарську помилку.

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase(), 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>

      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```


