---
id: 669840900adeba35e4ebb917
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

A link's text must be placed between the opening and closing tags of an anchor (`a`) element.

Here is an example of a link with the text `click here to go to freeCodeCamp.org`:

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

Додайте текст `link to cat pictures` до елемента посилання. This will become the link's text.

# --hints--

Елемент посилання (`a`) повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('a'));
```

Елемент посилання (`a`) повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.match(code, /<\/a\>/);
```

Текстом елемента посилання (`a`) має бути `link to cat pictures`. Make sure to put the link text between the anchor (`a`) element's opening tag and closing tag.

```js
assert.equal(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' '),
    'link to cat pictures'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```


