---
id: 6698457295f6403fa49050b8
title: Крок 25
challengeType: 0
dashedName: step-25
---

# --description--

Щоб підкреслити конкретне слово чи фразу, можна використати елемент `em`.

Відзначте слово `love` в елементі `figcaption`, обернувши його в елемент виділення `em`.

# --hints--

Елемент виділення `em` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('em'));
```

Елемент виділення `em` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.match(code, /<\/em\>/);
```

Ви або видалили елемент `figcaption`, або в ньому відсутній початковий чи кінцевий тег.

```js
assert.exists(document.querySelector('figcaption'));
assert.match(code, /<\/figcaption\>/);
```

Елемент виділення `em` повинен оточувати текст `love`. Ви або не написали текст, або маєте друкарську помилку.

```js
assert.equal(
  document.querySelector('figcaption > em')?.innerText?.toLowerCase(), 'love'
);
```

Текст `figcaption` повинен бути `Cats love lasagna`. Перевірте наявність друкарських помилок та необхідних пробілів навколо початкового та кінцевого теґів елемента `em`.

```js
assert.match(
  document
    .querySelector('figcaption')
    .innerText.replace(/\s+/gi, ' '),
    /cats love lasagna\.?/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```


