---
id: 66984ec64a3652522e81c678
title: Крок 35
challengeType: 0
dashedName: step-35
---

# --description--

Перетворіть наявний текст `freeCodeCamp.org` в посилання, вклавши його в елемент прив’язки (`a`). Атрибут `href` повинен бути встановленим на `https://www.freecodecamp.org`.

# --hints--

Елемент посилання (`a`) повинен бути вкладеним в межах елемента `p` всередині `footer`. Переконайтеся, що додали початковий та кінцевий теги для елемента посилання (`a`).

```js
assert.isNotEmpty(document.querySelectorAll('footer > p > a'));
```

Елемент посилання (`a`) повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert.exists(aElemClosingTags);
assert.lengthOf(aElemClosingTags, 3);
```

Елемент посилання (`a`) повинен мати атрибут `href` зі значенням `https://www.freecodecamp.org`. Можливо, ви пропустили атрибут/значення або маєте друкарську помилку.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(nestedAnchor.getAttribute('href').toLowerCase(), 'https://www.freecodecamp.org');
```

Текст посилання повинен бути `freeCodeCamp.org`. Ви або не написали текст, або маєте друкарську помилку.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' '),
    'freecodecamp.org'
);
```

Після вкладення елемента посилання (`a`) єдиним видимим вмістом елемента `p` в браузері повинен бути `No Copyright - freeCodeCamp.org`. Двічі перевірте текст, пробіли або пунктуацію обох елементів `p` та вкладеного елемента посилання.

```js
const pText = document.querySelectorAll('footer > p')[0].innerText.toLowerCase().replace(/\s+/g, ' ');
assert.match(pText, /^no copyright - freecodecamp\.org$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```


