---
id: 612e78af05201622d4bab8aa
title: Hatua ya 3
challengeType: 0
dashedName: step-3
---

# --description--

Time to start working on the piano. Create a `div` element within your `body` element with the `id` set to `piano`.

# --hints--

Unapaswa kuunda kipengele kipya cha `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Kipengele chako cha `div` kinapaswa kuwa ndani ya kipengele chako cha `body`.

```js
const div = document.querySelector('div');
assert(div?.parentElement?.localName === 'body');
```

Kipengee chako cha `div` kinapaswa kuwa na `id` ya `piano`.

```js
const div = document.querySelector('div');
assert(div?.getAttribute('id') === 'piano');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css

```
