---
id: 66bb6a9c2dd58b73cd759034
title: Створіть застосунок залікову книжку
challengeType: 25
dashedName: build-a-gradebook-app
---

# --description--

Виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Ви повинні мати функцію під назвою `getAverage`, яка приймає масив результатів тестів як параметр та повертає середній бал.
1. Ви повинні мати функцію під назвою `getGrade`, яка приймає бал студента як параметр та повертає рядок, що представляє літерну оцінку на основі балу. Ось бали та відповідні їм буквені оцінки:

    | Діапазон балів | Оцінка |
    | -------------- | ------ |
    | `100`          | `"A+"` |
    | `90 - 99`      | `"A"`  |
    | `80 - 89`      | `"B"`  |
    | `70 - 79`      | `"C"`  |
    | `60 - 69`      | `"D"`  |
    | `0 - 59`       | `"F"`  |

1. Ви повинні мати функцію під назвою `hasPassingGrade`, яка приймає оцінку як параметр та повертає `true` або `false` залежно від того, чи відповідає оцінка прохідному балу.
1. Функція `hasPassingGrade` має використовувати функцію `getGrade`, щоб отримати літерну оцінку та визначити, чи є оцінка прохідною. Прохідний бал — це все, що відрізняється від `"F"`.
1. Ви повинні мати функцію під назвою `studentMsg`, яка приймає як параметри масив балів та бал студента. Функція має повернути рядок у форматі:

    - `"Class average: average-goes-here. Your grade: grade-goes-here. You passed the course."`, якщо студент пройшов курс.
    - `"Class average: average-goes-here. Your grade: grade-goes-here. You failed the course."`, якщо студент завалив курс.

    Замініть `average-goes-here` на середнє значення загальної кількості балів, а `grade-goes-here` на оцінку студента. Використовуйте `getAverage`, щоб отримати середній бал, та `getGrade`, щоб отримати оцінку студента.

# --hints--

Ви повинні мати функцію під назвою `getAverage`.

```js
assert.isFunction(getAverage);
```

Функція `getAverage` має повернути число.

```js
assert.isNumber(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` має повернути `71.7`.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` має повернути `85.4`.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` має повернути `92.4`.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

Функція `getAverage` має повернути середній бал результатів тесту.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

Ви повинні мати функцію під назвою `getGrade`.

```js
assert.isFunction(getGrade);
```

Функція `getGrade` має повернути `"A+"`, якщо оцінка `100`.

```js
assert.strictEqual(getGrade(100), "A+");
```

Функція `getGrade` має повернути `"A"`, якщо оцінка між `90` та `99`.

```js
assert.strictEqual(getGrade(90), "A");
assert.strictEqual(getGrade(94), "A");
assert.strictEqual(getGrade(99), "A");
```

Функція `getGrade` має повернути `"B"`, якщо оцінка між `80` та `89`.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

Функція `getGrade` має повернути `"C"`, якщо оцінка між `70` та `79`.

```js
assert.strictEqual(getGrade(70), "C");
assert.strictEqual(getGrade(75), "C");
assert.strictEqual(getGrade(79), "C");
```

Функція `getGrade` має повернути `"D"`, якщо оцінка між `60` та `69`.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

Функція `getGrade` має повернути `"F"`, якщо оцінка між `0` та `59`.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
assert.strictEqual(getGrade(59), "F");
```


Ви повинні мати функцію під назвою `hasPassingGrade`.

```js
assert.isFunction(hasPassingGrade);
```

Функція `hasPassingGrade` має повернути булеве значення.

```js
assert.isBoolean(hasPassingGrade(100));
```

Функція `hasPassingGrade` має повернути `true`, якщо оцінка є `"A"`.

```js
assert.isTrue(hasPassingGrade(100));
```

Функція `hasPassingGrade` має повернути `false`, якщо оцінка є `"F"`.

```js
assert.isFalse(hasPassingGrade(53));
```

Функція `hasPassingGrade` має повернути `true` для всіх прохідних оцінок.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

Ви повинні мати функцію під назвою `studentMsg`.

```js
assert.isFunction(studentMsg);
```

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` має повернути таке повідомлення: `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` має повернути таке повідомлення: `"Class average: 50.8. Your grade: A+. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A+. You passed the course.");
```

Функція `studentMsg` має повернути правильне повідомлення на основі балів студента та середнього по класу.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function getAverage(scores) {
    let sum = 0;

    for (const score of scores) {
      sum += score;
    }

    return sum / scores.length;
  }

  function getGrade(score) {
    if (score === 100) {
      return "A+";
    } else if (score >= 90) {
      return "A";
    } else if (score >= 80) {
      return "B";
    } else if (score >= 70) {
      return "C";
    } else if (score >= 60) {
      return "D";
    } else {
      return "F";
    }
  }

  function hasPassingGrade(score) {
    return getGrade(score) !== "F";
  }

  function studentMsg(totalScores, studentScore) {
    let average = getAverage(totalScores);
    let grade = getGrade(studentScore);

    return `Class average: ${average}. Your grade: ${grade}. You ${
      hasPassingGrade(studentScore) ? "passed" : "failed"
    } the course.`;
  }
```
