---
id: 5f332a88dc25a0fd25c7687a
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

咖啡厅的名称是 `CAMPER CAFE`。  So, add an `h1` element within your `main` element. 给咖啡厅添加名称，用大写字母，使其突出显示。

# --hints--

你应该有一个开头的 `<h1>` 标签。

```js
assert.match(code, /<h1>/i);
```

你应该有一个闭合的 `</h1>` 标签。

```js
assert.match(code, /<\/h1>/i);
```

你不应该改变你现有的 `main` 元素。

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

你的 `h1` 元素应该嵌套在你的 `main` 元素中。

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

你的 `h1` 元素应该有文字 `CAMPER CAFE`。

```js
assert.equal(document.querySelector("h1")?.innerText.trim(), "CAMPER CAFE");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```

