---
id: 5f33294a6af5e9188dbdb8f3
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次工作坊中，你将通过创建一个咖啡馆菜单来练习 CSS（层叠样式表）的基础知识。

开始添加一些菜单内容。 在现有的 `main` 元素里面添加一个 `body` 元素。 它最终将包含有关咖啡馆提供的咖啡和甜点的价格信息。

# --hints--

你的代码应有一个开头的 `<main>` 标签。

```js
assert.match(code, /<main>/i);
```

你的代码应有一个闭合的 `</main>` 标签。

```js
assert.match(code, /<\/main>/i);
```

你不应该改变你的 `body` 元素。 确认你没有不小心删除结束标签。

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

你的 `main` 标签应该在你的 `body` 标签里面。

```js
const main = document.querySelector('main');
assert.equal(main.parentElement.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
