---
id: 5f332b23c2045fb843337579
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

为了让访客知道咖啡馆成立于 `2020`，在 `h1` 元素下方添加一个带有文本 `Est. 2020` 的 `p` 元素。

# --hints--

你应该有一个开头的 `<p>` 标签。

```js
assert.match(code, /<p>/i);
```

你应该有一个闭合的 `</p>` 标签。

```js
assert.match(code, /<\/p>/i);
```

你不应该改变你现有的 `h1` 元素。 确保你没有删除结束标签。

```js
assert.lengthOf(document.querySelectorAll('h1'),1);
```

你的 `p` 元素应该在 `h1` 元素的下面。

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.tagName, 'H1');
```

你的 `p` 元素应该有文本 `Est. 2020`。

```js
assert.equal(document.querySelector("p")?.innerText.trim(), "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```

