---
id: 5f344fbc22624a2976425065
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

在 `section` 元素中创建一个 `h2` 元素，并在它里面填写文字 `Coffee`。

# --hints--

你应该有一个开启的 `<h2>` 标签。

```js
assert.match(code, /<h2\s*>/i);
```

你应该有一个闭合的 `</h2>` 标签。

```js
assert.match(code, /<\/h2\s*>/i);
```

你不应该改变你现有的 `section` 元素。 确保你没有删除结束标签。

```js
assert.lengthOf(document.querySelectorAll('section'),1);
```

你的 `h2` 元素应该在 `section` 元素中。

```js
const h2 = document.querySelector('h2');
assert.equal(h2.parentElement.tagName, 'SECTION');
```

你的 `h2` 元素内应该有文本 `Coffee`。

```js
const h2 = document.querySelector('h2');
assert.equal(h2?.innerText.trim(), 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
