---
id: 5f344fad8bf01691e71a30eb
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

到目前为止，你对内容的展示和外观的控制有限。 为了改变这种情况，在 `head` 元素中创建一个 `style` 元素。

# --hints--

你应该有一个 `<style>` 开始标签。

```js
assert.match(code, /<style\s*>/i);
```

你应该有一个 `</style>` 结束标签。

```js
assert.match(code, /<\/style\s*>/);
```

你的 `style` 元素都应嵌套在 `head` 元素中。

```js
assert.match(code, /<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
