---
id: 5f3477ae34c1239cafe128be
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

你现在有三个样式完全相同的类型选择器。 你可以通过创建一个选择器列表将同一组样式添加到许多元素上。 每个选择器都用逗号分隔，如下所示：

```css
selector1, selector2 {
  property: value;
}
```

删除三个现有的类型选择器，并将它们替换为一个选择器列表，该列表使 `h1`、`h2` 和 `p` 元素的文本居中。

# --hints--

你应该对所有三个元素使用一个类型选择器，`h1, h2, p`。 请一定按照这个顺序。

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert.exists(hasSelector);
```

`style` 元素中应该只有一个选择器。

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert.equal(selectors.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
--fcc-editable-region--
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
--fcc-editable-region--
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

