---
id: 673b567e3ba535dda140d278
title: 투표 시스템 만들기
challengeType: 26
dashedName: build-a-voting-system
saveSubmissionToDB: true
---

# --description--

이 실습에서는 `Map`를 사용해 투표를 만들고 `Set`로 중복 투표를 방지하는 투표 시스템을 만드세요.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. `poll` 변수를 새 `Map` 객체로 초기화해야 합니다.

2. `option` 매개변수를 받는 `addOption` 함수를 만들어야 합니다.

3. `addOption` 함수에서는:

   - `option`이 투표에 아직 없으면, 빈 `Set`를 값으로 하여 `option`을 투표에 추가해 `voters`를 추적하세요. 그리고 메시지 `Option "<option>" added to the poll.`를 반환해야 합니다.

   - `option`이 이미 있으면 메시지 `Option "<option>" already exists.`를 반환해야 합니다.

   - 빈 옵션을 추가하려 하면 메시지 `Option cannot be empty.`를 반환해야 합니다.

4. `option` (투표할 옵션)과 `voterId` (투표자의 고유 ID) 두 매개변수를 받는 `vote` 함수를 만들어야 합니다.

5. `vote` 함수에서는:

   - `option`이 투표에 없으면 메시지 `Option "<option>" does not exist.`를 반환해야 합니다.

   - `option`이 있으면, `voterId`가 이미 이 `option`에 투표했는지 확인해야 합니다.

   - 투표자가 이미 투표했으면 메시지 `Voter <voterId> has already voted for "<option>".`를 반환해야 합니다.

   - 투표하지 않았다면 `voterId`를 `voters`의 `Set`에 추가해야 합니다. 그리고 메시지 `Voter <voterId> voted for "<option>".`를 반환해야 합니다.

6. `poll`에 최소 세 개의 옵션이 있어야 합니다.

7. `poll`에 최소 세 표가 있어야 합니다.

8. `displayResults` 함수를 만들어 다음 형식으로 투표 결과를 반환해야 합니다:

```js
Poll Results:
OptionA: N votes
OptionB: N votes
.
.

/*
sample output

Poll Results:
Turkey: 2 votes
Morocco: 1 votes
*/
```

# --hints--

`poll` 변수를 새 `Map` 객체로 초기화해야 합니다.

```js
assert.exists(poll);
assert.instanceOf(poll, Map);
```

`option` 매개변수를 받는 `addOption` 함수를 정의해야 합니다.

```js
assert.isFunction(addOption);
assert.match(
  addOption.toString(),
  /^function\s+addOption\s*\(\s*option\s*\)\s*\{|\(?\s*option\s*\)?\s*=>/
);
```

`option`와 `voterId` 두 매개변수를 받는 `vote` 함수를 정의해야 합니다.

```js
assert.isFunction(vote);
assert.match(vote.toString(), /\(\s*option\s*,\s*voterId\s*\)/);
```

투표 결과를 표시하는 `displayResults` 함수를 정의해야 합니다.

```js
assert.isFunction(displayResults);
```

`poll`에 최소 세 개의 옵션이 있어야 합니다.

```js
assert.isAtLeast(poll.size, 3);
```

`poll`에 최소 세 표가 있어야 합니다.

```js
let totalVotes = 0;
for (let voters of poll.values()) {
  totalVotes += voters.size;
}
assert.isAtLeast(totalVotes, 3);
```

각 투표 옵션이 `Set` 객체에 매핑되도록 해야 합니다.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();

  addOption('Turkey');
  assert.instanceOf(poll.get('Turkey'), Set);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`addOption("Egypt")`는 `Option "Egypt" added to the poll.`를 반환해야 합니다.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(addOption('Egypt'), 'Option "Egypt" added to the poll.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

빈 옵션을 추가하면 `"Option cannot be empty."`를 반환해야 합니다.

```js
assert.equal(addOption(''), 'Option cannot be empty.');
```

`Turkey`가 이미 추가되어 있으면 `addOption("Turkey")`는 `Option "Turkey" already exists.`를 반환해야 합니다.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Turkey');
  assert.equal(addOption('Turkey'), 'Option "Turkey" already exists.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`Malaysia`이 투표 옵션에 있으면 `vote("Malaysia", "traveler1")`는 `Voter traveler1 voted for "Malaysia".`를 반환해야 합니다.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Malaysia');
  assert.equal(
    vote('Malaysia', 'traveler1'),
    'Voter traveler1 voted for "Malaysia".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`vote`는 옵션의 투표자 `Set`를 업데이트해야 합니다.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Alaska');
  vote('Alaska', 'seal');
  assert.isTrue(poll.get('Alaska').has('seal'));
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`traveler1`가 `Algeria`에 다시 투표하려 하면 `vote("Algeria", "traveler1")`는 `Voter traveler1 has already voted for "Algeria".`를 반환해야 합니다.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Algeria');
  vote('Algeria', 'traveler1');
  assert.equal(
    vote('Algeria', 'traveler1'),
    'Voter traveler1 has already voted for "Algeria".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

중복 투표는 `Set` 크기를 늘리지 않아야 합니다.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  vote('Turkey', 'traveler1');
  const voterSet = poll.get('Turkey');
  const initialSize = voterSet.size;

  vote('Turkey', 'traveler1');
  assert.equal(voterSet.size, initialSize);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`Nigeria`이 투표 옵션에 없으면 `vote("Nigeria", "traveler2")`는 `Option "Nigeria" does not exist.`를 반환해야 합니다.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(
    vote('Nigeria', 'traveler2'),
    'Option "Nigeria" does not exist.'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

고유한 옵션은 여러 표를 받을 수 있어야 합니다.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Bali');
  assert.equal(vote('Bali', 'traveler1'), 'Voter traveler1 voted for "Bali".');
  assert.equal(vote('Bali', 'traveler2'), 'Voter traveler2 voted for "Bali".');
  assert.equal(poll.get('Bali').size, 2);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

옵션 `Turkey`이 두 표, 옵션 `Morocco`이 한 표, 옵션 `Spain`가 표를 받지 못하면 `displayResults()`는 `"Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes"`를 반환해야 합니다.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');
  addOption('Spain');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`displayResults()`는 올바른 형식으로 결과를 반환해야 합니다.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const poll = new Map();

function addOption(option) {
  if (!option || option.trim() === '') {
    return 'Option cannot be empty.';
  }
  if (!poll.has(option)) {
    poll.set(option, new Set());
    return `Option "${option}" added to the poll.`;
  } else {
    return `Option "${option}" already exists.`;
  }
}

function vote(option, voterId) {
  if (!poll.has(option)) {
    return `Option "${option}" does not exist.`;
  }
  const voters = poll.get(option);
  if (voters.has(voterId)) {
    return `Voter ${voterId} has already voted for "${option}".`;
  } else {
    voters.add(voterId);
    return `Voter ${voterId} voted for "${option}".`;
  }
}

function displayResults() {
  let results = 'Poll Results:\n';
  for (let [option, voters] of poll.entries()) {
    results += `${option}: ${voters.size} votes\n`;
  }
  return results.trim();
}

addOption('Turkey');
addOption('Morocco');
addOption('Spain');

vote('Turkey', 'traveler1');
vote('Turkey', 'traveler2');
vote('Morocco', 'traveler3');
```
