---
id: 68c362b379059c388d3874f2
title: デジタルペットゲームを作ろう
challengeType: 25
dashedName: lab-digital-pet-game
demoType: onClick
---

# --description--

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

アプリケーションは2つのビューを持つ必要があります。各ビューの要件は残りのユーザーストーリーに含まれています。

最初のビュー:

- 1つの文字列入力を受け付けるフォームを含む必要があります。
- ユーザーにペットの名前を尋ねる必要があります。
- ゲームを開始するための提出ボタンを含む必要があります。
- 提出時に非表示になる必要があります。

2番目のビュー:

- 最初のビューの提出時に表示される必要があります。
- 3つの`<button>`要素を含む必要があります。
- 各ボタン要素は次の3つのアクションのいずれかに対応する必要があります: `[EAT, PLAY, SLEEP]`。

  - `EAT`ボタンはペットの`Hunger`値を減らし、ペットの`Energy`値を増やす必要があります。
  - `PLAY`ボタンはペットの`Energy`値を減らし、ペットの`Happiness`値を増やす必要があります。
  - `SLEEP`ボタンはペットの`Hunger`値を増やし、ペットの`Energy`値を増やす必要があります。

- ペットの3つのプロパティを視覚的に追跡する必要があります。これらは`['Hunger', 'Happiness', 'Energy']`です。
- それぞれのプロパティを追跡する要素は以下を満たす必要があります:

  - そのプロパティに関連するテキストを含み、大文字小文字を区別しません。例: `hunger`、`Hunger`、`hUnGeR`はすべて有効です。
  - 他のプロパティに関連するテキストを含まないこと、つまり`Hunger`の追跡要素は`Happiness`のテキストを含んではいけません。
  - そのプロパティの現在の値を明確な数値で表すこと。

- ペットに利用可能な各`Action`は列挙子で囲まれる必要があります。
- 各`Action` `button`は、そのボタンのアクションに関連する次の`id`のいずれかを1つ含む必要があります: `['eat-action', 'play-action', 'sleep-action']`。
- ペットは次のムードを持ち、`PetMood`列挙子で囲まれる必要があります: `['Happy', 'Excited', 'Content', 'Sad', 'Tired', 'Sick', 'Hungry']`。

  - `Hungry`の場合、ペットは`Hunger > 70`である必要があります。
  - `Tired`の場合、ペットは`Energy < 30`である必要があります。
  - `Sad`の場合、ペットは`Happiness < 30`である必要があります。
  - `Excited`かつ`Happiness > 80`の場合、ペットは`Energy > 70`である必要があります。
  - `Happy`の場合、ペットは`Happiness > 60`である必要があります。
  - それ以外の場合、ペットは`Content`である必要があります。

- ペットは各`Mood`に結びついた複数の画像または絵文字を持つ必要があります。

  - これらの絵文字または画像は現在の`PetMood`に結びつき、`Record`を持つ必要があります。
  - 各`Record`エントリのキーは関連する`PetMood`列挙値である必要があります。

# --before-all--

```js
globalThis._PET_NAME = 'Fitzgerald';
globalThis._clock = __FakeTimers.install();
```

# --after-all--

```js
globalThis._clock.uninstall();
delete globalThis._clock;
delete globalThis._PET_NAME;
```

# --hints--

ページが読み込まれたとき、`form`要素が表示されている必要があります。

```js
assert.isNotNull(document.querySelector('form'), 'The form element does not exist at load.');
```

このフォームにはidが`pet-name`の入力フィールドが含まれている必要があります。

```js
assert.isNotNull(document.querySelector('#pet-name'));
```

このフォームには`button`が含まれている必要があります。

```js
assert.isNotNull(document.querySelector('button'));
```

入力に値が与えられ、ボタンが押されたとき、フォームは表示されなくなる必要があります。

```js
document.querySelector('#pet-name').value = _PET_NAME;
document.querySelector('button').click();
await _clock.tickAsync(500);

assert.isNull(document.querySelector('form'));
```

2番目のビューにはクラスが`pet-name`の要素が含まれている必要があります。

```js
const petNameEl = document.querySelector('.pet-name');
assert.exists(petNameEl);
```

2番目のビューには指定されたペットの名前が含まれている必要があります。

```js
const petNameEl = document.querySelector('.pet-name');
assert.strictEqual(petNameEl.textContent, _PET_NAME);
```

2番目のビューには3つの`button`が含まれている必要があります。

```js
const buttons = document.querySelectorAll('button');
assert.equal(buttons.length, 3);
```

idが`#eat-action`のボタンは1つだけ存在する必要があります。

```js
const eatButton = document.querySelectorAll('button#eat-action');
assert.equal(eatButton.length, 1);
```

idが`#play-action`のボタンは1つだけ存在する必要があります。

```js
const playButton = document.querySelectorAll('button#play-action');
assert.equal(playButton.length, 1);
```

idが`#sleep-action`のボタンは1つだけ存在する必要があります。

```js
const sleepButton = document.querySelectorAll('button#sleep-action');
assert.equal(sleepButton.length, 1);
```

クラスが`.stat`の要素は正確に3つ存在する必要があります。

```js
const statEls = document.querySelectorAll('.stat');
assert.equal(statEls.length, 3);
```

各ペットのステータスはクラスが`.stat-value`の要素を正確に1つ含む必要があります。

```js
const statEls = document.querySelectorAll('.stat');
for (const statEl of statEls) {
  assert.equal(statEl.querySelectorAll('.stat-value').length, 1);
}
```

各`.stat-value`要素は1つの数字を含む必要があります。

```js
const statValueEls = document.querySelectorAll('.stat-value');
for (const statEl of statValueEls) {
  assert.match(statEl.textContent, /-?\d+\.*\d*/);
}
```

各`.stat-value`の数字は`[0,100]`の範囲内（両端含む）である必要があります。

```js
const statValueEls = document.querySelectorAll('.stat-value');
for (const statEl of statValueEls) {
  const statValue = parseFloat(statEl.textContent);
  assert.isAtMost(statValue, 100);
  assert.isAtLeast(statValue, 0);
}
```

`.stat`の要素は1つだけで、テキスト`Hunger`を大文字小文字を区別せずに含む必要があります。

```js
const matches = Array.from(document.querySelectorAll('.stat')).filter(el => /hunger/i.test(el.textContent));
assert.equal(matches.length, 1);
```

`.stat`の要素は1つだけで、テキスト`Energy`を大文字小文字を区別せずに含む必要があります。

```js
const matches = Array.from(document.querySelectorAll('.stat')).filter(el => /energy/i.test(el.textContent));
assert.equal(matches.length, 1);
```

`.stat`の要素は1つだけで、テキスト`Happiness`を大文字小文字を区別せずに含む必要があります。

```js
const matches = Array.from(document.querySelectorAll('.stat')).filter(el => /happiness/i.test(el.textContent));
assert.equal(matches.length, 1);
```

`Hunger`のステータスは0から始まる必要があります。

```js
const hungerStatEl = Array.from(document.querySelectorAll('.stat')).find(el => /hunger/i.test(el.textContent));
const hungerValue = parseFloat(hungerStatEl.querySelector('.stat-value').textContent);
assert.equal(hungerValue, 0);
```

`Energy`のステータスは100から始まる必要があります。

```js
const energyStatEl = Array.from(document.querySelectorAll('.stat')).find(el => /energy/i.test(el.textContent));
const energyValue = parseFloat(energyStatEl.querySelector('.stat-value').textContent);
assert.equal(energyValue, 100);
```

`Happiness`のステータスは100から始まる必要があります。

```js
const happinessStatEl = Array.from(document.querySelectorAll('.stat')).find(el => /happiness/i.test(el.textContent));
const happinessValue = parseFloat(happinessStatEl.querySelector('.stat-value').textContent);
assert.equal(happinessValue, 100);
```

`Eat`ボタンをクリックすると`Hunger`の値が減り、`Energy`の値が増えます。`Hunger`が0または`Energy`が100の場合、その値は変わってはいけません。

```js
const statEls = () => Array.from(document.querySelectorAll('.stat'));
const hungerBefore = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyBefore = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
document.querySelector('#eat-action').click();
_clock.tick(50);
const hungerAfter = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyAfter = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
if (hungerBefore !== 0) assert.isBelow(hungerAfter, hungerBefore);
else assert.equal(hungerAfter, hungerBefore);
if (energyBefore !== 100) assert.isAbove(energyAfter, energyBefore);
else assert.equal(energyAfter, energyBefore);
```

`Play`ボタンをクリックすると`Energy`の値が減り、`Happiness`の値が増えます。`Energy`が0または`Happiness`が100の場合、その値は変わってはいけません。

```js
const statEls = () => Array.from(document.querySelectorAll('.stat'));
const energyBefore = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
const happinessBefore = parseFloat(statEls().find(el => /happiness/i.test(el.textContent)).querySelector('.stat-value').textContent);
document.querySelector('#play-action').click();
_clock.tick(50);
const energyAfter = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
const happinessAfter = parseFloat(statEls().find(el => /happiness/i.test(el.textContent)).querySelector('.stat-value').textContent);
if (energyBefore !== 0) assert.isBelow(energyAfter, energyBefore);
else assert.equal(energyAfter, energyBefore);
if (happinessBefore !== 100) assert.isAbove(happinessAfter, happinessBefore);
else assert.equal(happinessAfter, happinessBefore);
```

`Sleep`ボタンをクリックすると`Hunger`の値が増え、`Energy`の値が増えます。`Hunger`が100または`Energy`が100の場合、その値は変わってはいけません。

```js
const statEls = () => Array.from(document.querySelectorAll('.stat'));
const hungerBefore = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyBefore = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
document.querySelector('#sleep-action').click();
_clock.tick(50);
const hungerAfter = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyAfter = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
if (hungerBefore !== 100) assert.isAbove(hungerAfter, hungerBefore);
else assert.equal(hungerAfter, hungerBefore);
if (energyBefore !== 100) assert.isAbove(energyAfter, energyBefore);
else assert.equal(energyAfter, energyBefore);
```

`Hunger`のステータスは長時間の放置後に100になる必要があります。

```js
await _clock.tickAsync(1_200_000);

const hungerStatEl = Array.from(document.querySelectorAll('.stat')).find(el => /hunger/i.test(el.textContent));
const hungerValue = parseFloat(hungerStatEl.querySelector('.stat-value').textContent);
assert.equal(hungerValue, 100);
```

`Energy`のステータスは長時間の放置後に100になる必要があります。

```js
const energyStatEl = Array.from(document.querySelectorAll('.stat')).find(el => /energy/i.test(el.textContent));
const energyValue = parseFloat(energyStatEl.querySelector('.stat-value').textContent);
assert.equal(energyValue, 100);
```

`Happiness`のステータスは長時間の放置後に0になる必要があります。

```js
const happinessStatEl = Array.from(document.querySelectorAll('.stat')).find(el => /happiness/i.test(el.textContent));
const happinessValue = parseFloat(happinessStatEl.querySelector('.stat-value').textContent);
assert.equal(happinessValue, 0);
```

`Eat`ボタンをクリックすると`Hunger`の値が減り、`Energy`の値が増えます。`Hunger`が0または`Energy`が100の場合、その値は変わってはいけません。

```js
const statEls = () => Array.from(document.querySelectorAll('.stat'));
const hungerBefore = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyBefore = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
document.querySelector('#eat-action').click();
_clock.tick(50);
const hungerAfter = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyAfter = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
if (hungerBefore !== 0) assert.isBelow(hungerAfter, hungerBefore);
else assert.equal(hungerAfter, hungerBefore);
if (energyBefore !== 100) assert.isAbove(energyAfter, energyBefore);
else assert.equal(energyAfter, energyBefore);
```

`Play`ボタンをクリックすると`Energy`の値が減り、`Happiness`の値が増えます。`Energy`が0または`Happiness`が100の場合、その値は変わってはいけません。

```js
const statEls = () => Array.from(document.querySelectorAll('.stat'));
const energyBefore = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
const happinessBefore = parseFloat(statEls().find(el => /happiness/i.test(el.textContent)).querySelector('.stat-value').textContent);
document.querySelector('#play-action').click();
_clock.tick(50);
const energyAfter = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
const happinessAfter = parseFloat(statEls().find(el => /happiness/i.test(el.textContent)).querySelector('.stat-value').textContent);
if (energyBefore !== 0) assert.isBelow(energyAfter, energyBefore);
else assert.equal(energyAfter, energyBefore);
if (happinessBefore !== 100) assert.isAbove(happinessAfter, happinessBefore);
else assert.equal(happinessAfter, happinessBefore);
```

`Sleep`ボタンをクリックすると`Hunger`の値が増え、`Energy`の値が増えます。`Hunger`が100または`Energy`が100の場合、その値は変わってはいけません。

```js
const statEls = () => Array.from(document.querySelectorAll('.stat'));
const hungerBefore = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyBefore = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
document.querySelector('#sleep-action').click();
_clock.tick(50);
const hungerAfter = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyAfter = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
if (hungerBefore !== 100) assert.isAbove(hungerAfter, hungerBefore);
else assert.equal(hungerAfter, hungerBefore);
if (energyBefore !== 100) assert.isAbove(energyAfter, energyBefore);
else assert.equal(energyAfter, energyBefore);
```

`PetMood`の列挙子が存在する必要があります。

```js
assert.match(code, /\s*enum PetMood\s*/);
```

`PetMood`の列挙子は`HAPPY`を含む必要があります。

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*HAPPY[\s*\w+,\s*]*}/);
```

`PetMood`の列挙子は`EXCITED`を含む必要があります。

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*EXCITED[\s*\w+,\s*]*}/);
```

`PetMood`の列挙子は`CONTENT`を含む必要があります。

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*CONTENT[\s*\w+,\s*]*}/);
```

`PetMood`の列挙子は`SAD`を含む必要があります。

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*SAD[\s*\w+,\s*]*}/);
```

`PetMood`の列挙子は`TIRED`を含む必要があります。

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*TIRED[\s*\w+,\s*]*}/);
```

`PetMood`の列挙子は`SICK`を含む必要があります。

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*SICK[\s*\w+,\s*]*}/);
```

`PetMood`の列挙子は`HUNGRY`を含む必要があります。

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*HUNGRY[\s*\w+,\s*]*}/);
```

`Record<PetMood, string`が存在し、各`PetMood`に値をマッピングする必要があります。

```js
assert.match(code, /\s*(?:const|let|var)\s+\w+:\s+Record<PetMood,\s*string>/);
```

`PetMood.HAPPY`にマッピングされた値が存在する必要があります。

```js
assert.match(code, /\[PetMood\.HAPPY\]\s*:/);
```

`PetMood.EXCITED`にマッピングされた値が存在する必要があります。

```js
assert.match(code, /\[PetMood\.EXCITED\]\s*:/);
```

`PetMood.CONTENT`にマッピングされた値が存在する必要があります。

```js
assert.match(code, /\[PetMood\.CONTENT\]\s*:/);
```

`PetMood.SAD`にマッピングされた値が存在する必要があります。

```js
assert.match(code, /\[PetMood\.SAD\]\s*:/);
```

`PetMood.TIRED`にマッピングされた値が存在する必要があります。

```js
assert.match(code, /\[PetMood\.TIRED\]\s*:/);
```

`PetMood.SICK`にマッピングされた値が存在する必要があります。

```js
assert.match(code, /\[PetMood\.SICK\]\s*:/);
```

`PetMood.HUNGRY`にマッピングされた値が存在する必要があります。

```js
assert.match(code, /\[PetMood\.HUNGRY\]\s*:/);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="utf-8"/>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Digital Pet Game</title>

    <link rel="preconnect" href="https://fonts.googleapis.com">
    <link rel="preconnect" href="https://fonts.gstatic.com" crossorigin>
    <link href="https://fonts.googleapis.com/css2?family=Tektur:wght@400..900&display=swap" rel="stylesheet">

        <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.tsx"
    ></script>
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { PetGame } from './index.tsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<PetGame />);
    </script>
  </body>
</html>
```

```css
:root {
  /* Backgrounds */
  --bg-gradient-start: #dbeafe;
  --bg-gradient-mid: #faf5ff;
  --bg-gradient-end: #fce7f3;

  /* Text */
  --text-primary: #1f2937;
  --text-muted: #6b7280;
  --text-muted-2: #9ca3af;

  /* Card / surface */
  --card-bg-translucent: rgba(255, 255, 255, 0.6);
  --card-bg-solid: rgba(255, 255, 255, 0.8);

  /* Accent / brand */
  --accent-purple: #8b5cf6;
  --accent-purple-dark: #7c3aed;

  /* Actions */
  --action-feed: #10b981;
  --action-feed-dark: #059669;
  --action-play: #3b82f6;
  --action-play-dark: #2563eb;
  --action-sleep: #8b5cf6;
  --action-sleep-dark: #7c3aed;
  --action-heal: #ec4899;
  --action-heal-dark: #db2777;

  /* Stats */
  --stat-high: #10b981;
  --stat-medium: #f59e0b;
  --stat-low: #ef4444;

  /* UI greys */
  --ui-bg-light: #e5e7eb;
  --ui-fill: #8b5cf6; /* used for exp fill as accent */

  /* Shadows */
  --shadow-strong: rgba(0, 0, 0, 0.25);
  --shadow-soft: rgba(0, 0, 0, 0.1);
}

html {
  /*color-scheme: light dark;*/
  font-family: system-ui;
  display: flex;
  justify-content: center;
  font-family:
    -apple-system, BlinkMacSystemFont, "Segoe UI", "Roboto", "Oxygen", "Ubuntu",
    "Cantarell", "Fira Sans", "Droid Sans", "Helvetica Neue", sans-serif;
  -webkit-font-smoothing: antialiased;
  -moz-osx-font-smoothing: grayscale;
  background: linear-gradient(135deg, #dbeafe 0%, #faf5ff 50%, #fce7f3 100%);
  min-height: 100vh;
  padding: 1rem;
}

header {
  text-align: center;
  margin-bottom: 2rem;
}

body {
  max-width: 400px;
}

.base-container {
  background: rgba(255, 255, 255, 0.6);
  -webkit-backdrop-filter: blur(10px);
  backdrop-filter: blur(10px);
  border-radius: 1.5rem;
  padding: 2rem;
  box-shadow: 0 25px 50px -12px rgba(0, 0, 0, 0.25);
  .pet-name {
    text-align: center;
  }
}

.game-container {
  margin-bottom: 1.5rem;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
}

/* Stats Grid */
.stats-grid {
  display: grid;
  grid-template-columns: repeat(2, 1fr);
  gap: 1rem;
  margin-bottom: 1.5rem;
}

.stat-bar {
  background: rgba(255, 255, 255, 0.8);
  -webkit-backdrop-filter: blur(10px);
  backdrop-filter: blur(10px);
  border-radius: 0.75rem;
  padding: 1rem;
  box-shadow: 0 1px 3px 0 rgba(0, 0, 0, 0.1);
  transition: all 0.2s ease;
}

.stat-bar:hover {
  box-shadow: 0 4px 6px -1px rgba(0, 0, 0, 0.1);
}

.stat-header {
  display: flex;
  align-items: center;
  justify-content: space-between;
  margin-bottom: 0.5rem;
}

.stat-label {
  display: flex;
  align-items: center;
  gap: 0.5rem;
}

.stat-icon {
  font-size: 1.25rem;
}

.stat-name {
  font-weight: 500;
  color: #374151;
}

.stat-value {
  font-weight: bold;
  color: #1f2937;
}

.stat-progress {
  width: 100%;
  height: 0.75rem;
  background: #e5e7eb;
  border-radius: 9999px;
  overflow: hidden;
}

.stat-fill {
  height: 100%;
  transition: width 0.5s ease-out;
  border-radius: 9999px;
}

.stat-fill.high {
  background: #10b981;
}

.stat-fill.medium {
  background: #f59e0b;
}

.stat-fill.low {
  background: #ef4444;
}

.pet-buttons {
  width: auto;
  display: flex;
  flex-direction: row;
  gap: 1rem;
}

.pet-buttons .pet-button {
  min-width: 80px;
  background: #c2214a;
  color: white;
  border-radius: 13%;
  border-style: solid;
  border-color: #870b04;
  box-shadow: rgba(120, 25, 23, 0.2) 2px 2px 4px 2px;
  &:active {
    transform: scale(0.95);
    box-shadow: 0 2px 4px rgba(0, 0, 0, 0.15) inset;
  }
  text-align: center;
  cursor: pointer;
  padding: 0.5rem 1rem;
}

.pet-sprite {
  font-size: xxx-large;
  text-align: center;
}

.data-management * {
  margin: 0 0.5rem 0 0;
}

.info-panel {
  text-align: center;
}
.info-panel #hud,
.info-panel .start-questions {
  display: flex;
  align-items: center;
  flex-direction: column;
  gap: 1rem;
}

.info-panel #hud p {
  margin: 0rem;
}

.info-panel .start-questions div * {
  margin: 0 0.5rem 0 0;
}

#pet-fact {
  color: var(--text-muted-2);
  margin: 0 0 0.5rem 0;
}
#set-name-btn {
  display: inline-block;
  padding: 14px 28px;
  font-size: 1.15rem;
  font-weight: 600;
  background: linear-gradient(135deg, #6b8cff, #5ec2ff);
  color: #fff;
  border: none;
  border-radius: 10px;
  box-shadow: 0 6px 18px rgba(92, 132, 255, 0.25);
  cursor: pointer;
  transition:
    transform 0.08s ease,
    box-shadow 0.12s ease,
    opacity 0.12s;
}

#set-name-btn:hover {
  opacity: 0.98;
}

#set-name-btn:active,
#set-name-btn.pressed {
  transform: translateY(2px) scale(0.995);
  box-shadow: 0 3px 8px rgba(92, 132, 255, 0.18);
}

#set-name-btn:focus {
  outline: 3px solid rgba(94, 190, 255, 0.28);
  outline-offset: 2px;
}
```

```tsx
const { useState } = React;

export const PetGame = () => {};
```

# --solutions--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="utf-8"/>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Digital Pet Game</title>

    <link rel="preconnect" href="https://fonts.googleapis.com">
    <link rel="preconnect" href="https://fonts.gstatic.com" crossorigin>
    <link href="https://fonts.googleapis.com/css2?family=Tektur:wght@400..900&display=swap" rel="stylesheet">
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.tsx"
    ></script>
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { PetGame } from './index.tsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<PetGame />);
    </script>
  </body>
</html>
```

```css
:root {
  /* Backgrounds */
  --bg-gradient-start: #dbeafe;
  --bg-gradient-mid: #faf5ff;
  --bg-gradient-end: #fce7f3;

  /* Text */
  --text-primary: #1f2937;
  --text-muted: #6b7280;
  --text-muted-2: #9ca3af;

  /* Card / surface */
  --card-bg-translucent: rgba(255, 255, 255, 0.6);
  --card-bg-solid: rgba(255, 255, 255, 0.8);

  /* Accent / brand */
  --accent-purple: #8b5cf6;
  --accent-purple-dark: #7c3aed;

  /* Actions */
  --action-feed: #10b981;
  --action-feed-dark: #059669;
  --action-play: #3b82f6;
  --action-play-dark: #2563eb;
  --action-sleep: #8b5cf6;
  --action-sleep-dark: #7c3aed;
  --action-heal: #ec4899;
  --action-heal-dark: #db2777;

  /* Stats */
  --stat-high: #10b981;
  --stat-medium: #f59e0b;
  --stat-low: #ef4444;

  /* UI greys */
  --ui-bg-light: #e5e7eb;
  --ui-fill: #8b5cf6; /* used for exp fill as accent */

  /* Shadows */
  --shadow-strong: rgba(0, 0, 0, 0.25);
  --shadow-soft: rgba(0, 0, 0, 0.1);
}

html {
  /*color-scheme: light dark;*/
  font-family: system-ui;
  display: flex;
  justify-content: center;
  font-family:
    -apple-system, BlinkMacSystemFont, "Segoe UI", "Roboto", "Oxygen", "Ubuntu",
    "Cantarell", "Fira Sans", "Droid Sans", "Helvetica Neue", sans-serif;
  -webkit-font-smoothing: antialiased;
  -moz-osx-font-smoothing: grayscale;
  background: linear-gradient(135deg, #dbeafe 0%, #faf5ff 50%, #fce7f3 100%);
  min-height: 100vh;
  padding: 1rem;
}

header {
  text-align: center;
  margin-bottom: 2rem;
}

body {
  max-width: 400px;
}

.base-container {
  background: rgba(255, 255, 255, 0.6);
  -webkit-backdrop-filter: blur(10px);
  backdrop-filter: blur(10px);
  border-radius: 1.5rem;
  padding: 2rem;
  box-shadow: 0 25px 50px -12px rgba(0, 0, 0, 0.25);
  .pet-name {
    text-align: center;
  }
}

.game-container {
  margin-bottom: 1.5rem;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
}

/* Stats Grid */
.stats-grid {
  display: grid;
  grid-template-columns: repeat(2, 1fr);
  gap: 1rem;
  margin-bottom: 1.5rem;
}

.stat-bar {
  background: rgba(255, 255, 255, 0.8);
  -webkit-backdrop-filter: blur(10px);
  backdrop-filter: blur(10px);
  border-radius: 0.75rem;
  padding: 1rem;
  box-shadow: 0 1px 3px 0 rgba(0, 0, 0, 0.1);
  transition: all 0.2s ease;
}

.stat-bar:hover {
  box-shadow: 0 4px 6px -1px rgba(0, 0, 0, 0.1);
}

.stat-header {
  display: flex;
  align-items: center;
  justify-content: space-between;
  margin-bottom: 0.5rem;
}

.stat-label {
  display: flex;
  align-items: center;
  gap: 0.5rem;
}

.stat-icon {
  font-size: 1.25rem;
}

.stat-name {
  font-weight: 500;
  color: #374151;
}

.stat-value {
  font-weight: bold;
  color: #1f2937;
}

.stat-progress {
  width: 100%;
  height: 0.75rem;
  background: #e5e7eb;
  border-radius: 9999px;
  overflow: hidden;
}

.stat-fill {
  height: 100%;
  transition: width 0.5s ease-out;
  border-radius: 9999px;
}

.stat-fill.high {
  background: #10b981;
}

.stat-fill.medium {
  background: #f59e0b;
}

.stat-fill.low {
  background: #ef4444;
}

.pet-buttons {
  width: auto;
  display: flex;
  flex-direction: row;
  gap: 1rem;
}

.pet-buttons .pet-button {
  min-width: 80px;
  background: #c2214a;
  color: white;
  border-radius: 13%;
  border-style: solid;
  border-color: #870b04;
  box-shadow: rgba(120, 25, 23, 0.2) 2px 2px 4px 2px;
  &:active {
    transform: scale(0.95);
    box-shadow: 0 2px 4px rgba(0, 0, 0, 0.15) inset;
  }
  text-align: center;
  cursor: pointer;
  padding: 0.5rem 1rem;
}

.pet-sprite {
  font-size: xxx-large;
  text-align: center;
}

.data-management * {
  margin: 0 0.5rem 0 0;
}

.info-panel {
  text-align: center;
}
.info-panel #hud,
.info-panel .start-questions {
  display: flex;
  align-items: center;
  flex-direction: column;
  gap: 1rem;
}

.info-panel #hud p {
  margin: 0rem;
}

.info-panel .start-questions div * {
  margin: 0 0.5rem 0 0;
}

#pet-fact {
  color: var(--text-muted-2);
  margin: 0 0 0.5rem 0;
}
#set-name-btn {
  display: inline-block;
  padding: 14px 28px;
  font-size: 1.15rem;
  font-weight: 600;
  background: linear-gradient(135deg, #6b8cff, #5ec2ff);
  color: #fff;
  border: none;
  border-radius: 10px;
  box-shadow: 0 6px 18px rgba(92, 132, 255, 0.25);
  cursor: pointer;
  transition:
    transform 0.08s ease,
    box-shadow 0.12s ease,
    opacity 0.12s;
}

#set-name-btn:hover {
  opacity: 0.98;
}

#set-name-btn:active,
#set-name-btn.pressed {
  transform: translateY(2px) scale(0.995);
  box-shadow: 0 3px 8px rgba(92, 132, 255, 0.18);
}

#set-name-btn:focus {
  outline: 3px solid rgba(94, 190, 255, 0.28);
  outline-offset: 2px;
}
```

```tsx
const { useState, useEffect } = React;

interface StatBarProps {
  label: string;
  value: number;
  icon: string;
  reverse?: boolean;
}

function getStatColorClass(value: number): string {
  if (value >= 70) return "high";
  if (value >= 40) return "medium";
  return "low";
}

function StatBar({ label, value, icon, reverse = false }: StatBarProps) {
  const colorClass = getStatColorClass(!reverse ? value : 100 - value);

  return (
    <div className="stat-bar stat">
      <div className="stat-header">
        <div className="stat-label">
          <span className="stat-icon">{icon}</span>
          <span className="stat-name">{label}</span>
        </div>
        <span className="stat-value">{Math.round(value)}%</span>
      </div>

      <div className="stat-progress">
        <div
          className={`stat-fill ${colorClass}`}
          style={{ width: `${value}%` }}
        />
      </div>
    </div>
  );
}

enum PetAction {
  NONE,
  EAT,
  PLAY,
  SLEEP,
}

interface Pet {
  name: string;
  happiness: number;
  hunger: number;
  energy: number;
  species: string;
}

enum PetMood {
  HAPPY,
  EXCITED,
  CONTENT,
  SAD,
  TIRED,
  SICK,
  HUNGRY,
}

export const STAT_DECAY_RATES = {
  hunger: 10,
  happiness: 5,
  energy: 5,
};

const UPDATE_INTERVAL = 30 * 1000; // 30 seconds

export function calculatePetMood(pt: Pet): PetMood {
  const { hunger, happiness, energy } = pt;

  if (hunger > 70) return PetMood.HUNGRY;
  if (energy < 30) return PetMood.TIRED;
  if (happiness < 30) return PetMood.SAD;
  if (happiness > 80 && energy > 70) return PetMood.EXCITED;
  if (happiness > 60) return PetMood.HAPPY;

  return PetMood.CONTENT;
}

const moodEmojiMap: Record<PetMood, string> = {
  [PetMood.HAPPY]: "😺", // grinning cat
  [PetMood.EXCITED]: "😻", // heart-eyes cat
  [PetMood.CONTENT]: "😸", // smiling cat
  [PetMood.SAD]: "😿", // crying cat
  [PetMood.TIRED]: "😽", // kissing cat (closest to sleepy)
  [PetMood.SICK]: "🙀", // weary cat (looks unwell)
  [PetMood.HUNGRY]: "😹", // joy cat (closest to hungry/funny)
};

interface usePetProps {
  isGameStarted: boolean;
  setGameStarted: React.Dispatch<React.SetStateAction<boolean>>;
}
// hook where all pet logic is handled
function usePet({ isGameStarted }: usePetProps) {
  let [pet, setPet] = useState<Pet>({
    name: "",
    happiness: 100,
    hunger: 0,
    energy: 100,
    species: "Cat",
  });
  // Auto-decay stats over time
  useEffect(() => {
    if (!isGameStarted) return;

    const interval = setInterval(() => {
      setPet((pet) => ({
        ...pet,
        happiness: Math.max(pet.happiness - 5, 0),
        hunger: Math.min(pet.hunger + 5, 100),
      }));
    }, UPDATE_INTERVAL);

    return () => clearInterval(interval);
  }, [isGameStarted]);

  function doAction(action: PetAction) {
    switch (action) {
      case PetAction.EAT:
        feedPet();
        break;
      case PetAction.SLEEP:
        restPet();
        break;
      case PetAction.PLAY:
        playWithPet();
        break;
    }
  }

  function feedPet() {
    setPet((pet) => ({
      ...pet,
      hunger: Math.max(pet.hunger - STAT_DECAY_RATES.hunger, 0),
      energy: Math.min(pet.energy + STAT_DECAY_RATES.energy, 100),
    }));
  }

  function playWithPet() {
    setPet((pet) => ({
      ...pet,
      energy: Math.max(pet.energy - STAT_DECAY_RATES.energy, 0),
      happiness: Math.min(pet.happiness + STAT_DECAY_RATES.happiness, 100),
    }));
  }

  function restPet() {
    setPet((pet) => ({
      ...pet,
      hunger: Math.min(pet.hunger + STAT_DECAY_RATES.hunger, 100),
      energy: Math.min(pet.energy + STAT_DECAY_RATES.energy, 100),
    }));
  }

  return {
    pet,
    doAction,
    setName: (name: string) => setPet({ ...pet, name }),
  };
}

export function PetGame() {
  const [isGameStarted, setGameStarted] = useState(false);

  let { pet, doAction, setName } = usePet({
    isGameStarted,
    setGameStarted,
  });

  const [fact, setFact] = useState<string | null>("");

  useEffect(() => {
    fetch("https://cat-facts-api.freecodecamp.rocks/api/catfacts/random").then(
      res => res.json()
    ).then(
      fact => setFact(fact)
    ).catch(err => {
      setFact("Sorry, we're not able to retrieve your cat fact right now!");
    });
  }, []);

  function startGame() {
    const petName = (document.getElementById("pet-name") as HTMLInputElement)
      .value;
    if (!petName) return; //do not let user to start game if no name selected
    setName(petName);
    setGameStarted(true);
  }

  return (
    <main>
      <header>
        <h1>Digital Pet Game</h1>
        <p>Take care of your virtual companion!</p>
      </header>

      {isGameStarted && (
        <section className="base-container game-container">
          <div className="pet-screen">
            <div className="pet-sprite">
              {moodEmojiMap[calculatePetMood(pet)]}
            </div>
            <h2 className="pet-name">{pet.name}</h2>
          </div>
          <div className="pet-buttons">
            <button
              onClick={() => doAction(PetAction.EAT)}
              className="pet-button pet-buttons-left"
              id="eat-action"
            >
              EAT
            </button>
            <button
              onClick={() => doAction(PetAction.PLAY)}
              className="pet-button pet-buttons-center"
              id="play-action"
            >
              PLAY
            </button>
            <button
              onClick={() => doAction(PetAction.SLEEP)}
              className="pet-button pet-buttons-right"
              id="sleep-action"
            >
              SLEEP
            </button>
          </div>
        </section>
      )}

      {isGameStarted && (
        <section className="stats-grid">
          <StatBar label="Hunger" value={pet.hunger} icon="🍽️" reverse />
          <StatBar label="Happiness" value={pet.happiness} icon="😊" />
          <StatBar label="Energy" value={pet.energy} icon="⚡" />
        </section>
      )}

      <section className="base-container info-panel">
        {!isGameStarted ? (
          <form className="start-questions">
            <label htmlFor="pet-name">What is your pet's name?</label>
            <input
              id="pet-name"
              name="pet-name"
              required={true}
              pattern="[A-Za-z0-9]{1,20}"
            />

            <button id="set-name-btn" onClick={startGame}>
              Start Game
            </button>
          </form>
        ) : (
          <div id="hud">
            <p id="pet-species">Species: {pet.species}</p>
            <p id="pet-fact">
              <b>Pet Fact:</b> {fact}
            </p>
          </div>
        )}
      </section>
    </main>
  );
}
```
