---
id: 68c15a2a1f0e70ca7154aba5
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 워크숍에서는 직원들의 급여 추적 시스템을 만들 것입니다.

`Employee`라는 이름의 클래스를 만드는 것부터 시작하세요. 그 안에 `__init__` 메서드를 `self`, `name`, 그리고 `level` 매개변수와 함께 만드세요. `__init__` 메서드 내에서 `name`와 `level`를 같은 이름의 인스턴스 속성에 할당하세요.

# --hints--

`Employee`라는 이름의 클래스가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Employee")`)) })
```

`Employee` 클래스에는 `__init__` 메서드가 세 개의 매개변수 `self`, `name`, `level`와 함께 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_args("self, name, level")`)) })
```

`__init__` 메서드는 `self.name`를 `name`에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

`__init__` 메서드는 `self.level`를 `level`에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
