---
id: 68c16924498728d259e4781c
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

이제 `Employee` 클래스를 인스턴스화하면서 문자열 `Charlie Brown`과 `trainee`를 전달하세요. 인스턴스를 `charlie_brown`라는 변수에 할당하세요.

# --hints--

`charlie_brown`라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("charlie_brown")`)) })
```

`Employee('Charlie Brown', 'trainee')`를 `charlie_brown`에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("charlie_brown").is_equivalent("charlie_brown = Employee('Charlie Brown', 'trainee')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level


--fcc-editable-region--
```
