---
id: 68c18b1d0fce081f85312d97
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

이전 수업에서, 단일 밑줄 접두사가 붙은 속성은 관례상 내부용임을 배웠습니다.

`name`와 `level` 속성 둘 다 `_name`와 `_level`로 수정하세요. 이 속성들은 클래스 외부에서 수정해서는 안 되기 때문입니다.

이것이 클래스 외부에서 속성에 접근하거나 수정하는 것을 막지는 않는다는 점을 참고하세요. 또한, Python에서는 항상 이중 밑줄 접두사가 붙은 비공개 속성에도 접근할 방법이 있습니다.

# --hints--

`__init__` 메서드는 `self._name`를 `name`로 설정해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._name = name")`)) })
```

`__init__` 메서드는 `self._level`를 `level`로 설정해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level

charlie_brown = Employee('Charlie Brown', 'trainee')
--fcc-editable-region--
```
