---
id: 68c28d659d345270d4a61084
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

`__repr__` 메서드는 객체를 생성하는 데 사용할 수 있는 문자열 표현을 반환해야 하는 특별한 메서드입니다.

예를 들어, `__repr__`의 `Employee('Charlie', 'developer')` 메서드는 객체를 생성하는 데 사용된 문자열과 같은 `Employee('Charlie', 'developer')` 문자열을 반환해야 합니다.

`Employee` 클래스에 `__repr__` 메서드를 `self` 매개변수와 함께 추가하고, 객체를 생성하는 데 사용할 수 있는 문자열을 반환하게 하세요.

# --hints--

`Employee` 클래스에는 `__repr__` 매개변수를 가진 `self` 메서드가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__repr__").has_args("self")`)) })
```

`__repr__` 메서드는 객체를 생성하는 데 사용된 문자열을 반환해야 합니다.

```js
({ test: () => runPython(`
  empl = Employee("Frank", "dreamer")
  assert repr(empl) == 'Employee("Frank", "dreamer")' or repr(empl) == "Employee('Frank', 'dreamer')"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
