---
id: 68c2d31a86d5672ca49d4521
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

이제 `__init__` 메서드에 유효성 검사를 추가할 시간입니다. 메서드 시작 부분에 `if` 문을 만들어 `name` 또는 `level` 중 하나라도 `str` 인스턴스가 아닌지 확인하세요.

`if` 문 안에서 메시지 `TypeError`와 함께 `'name' and 'level' attribute must be of type 'str'.`를 발생시키세요.

# --hints--

`if` 메서드에 `__init__` 문이 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0]`)) })
```

`if` 문은 `name`과 `level` 중 하나라도 `str` 인스턴스가 아닌지 확인해야 합니다.

```js
({ test: () => runPython(`
  cond = _Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0].find_conditions()[0]
  conditions = [
    "not (isinstance(name, str) and isinstance(level, str))",
    "not isinstance(name, str) or not isinstance(level, str)",
    "not (isinstance(level, str) and isinstance(name, str))",
    "not isinstance(level, str) or not isinstance(name, str)",
  ]
  assert any(cond.is_equivalent(c) for c in conditions)
`) })
```

`if` 문은 메시지 `TypeError`와 함께 `'name' and 'level' attribute must be of type 'str'.`를 발생시켜야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError("\\'name\\' and \\'level\\' attribute must be of type \\'str\\'.")')`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
--fcc-editable-region--
    def __init__(self, name, level):
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
