---
id: 68c28cfab52fef6913d18166
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

이제 마지막 두 개의 `print` 호출을 제거하세요.

# --hints--

코드에 `print(charlie_brown.name)`이 없어야 합니다.

```js
({ test: () => assert(runPython(`not _Node(_code).has_stmt("print(charlie_brown.name)")`)) })
```

코드에 `print(charlie_brown.level)`가 없어야 합니다.

```js
({ test: () => assert(runPython(`not _Node(_code).has_stmt("print(charlie_brown.level)")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self.name}: {self.level}'

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
print(charlie_brown.name)
print(charlie_brown.level)
--fcc-editable-region--
```
