---
id: 68c28b7cc3ab8f4daf144b5e
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

이전 단계에서 한 작업을 따라 `_level` 속성에 대한 getter를 만드세요.

# --hints--

`Employee` 클래스에는 `level` 매개변수를 가진 `self` 메서드가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_args("self")`)) })
```

`level` 메서드는 `@property`로 데코레이트되어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_decorators("property")`)) })
```

`level` 메서드는 `self._level`을 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_return("self._level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    
--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)

```
