---
id: 68c28c6517e7de5e85f8ebe5
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

이제 `_name`와 `_level` 속성에 대한 getter가 있으니, `__str__`가 반환하는 문자열을 `self.name`와 `self.level`를 사용하도록 업데이트하세요. 이렇게 하면 속성에 직접 접근하는 대신 getter를 호출하게 됩니다.

# --hints--

`__str__` 메서드는 `f'{self.name}: {self.level}'`를 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_return("f'{self.name}: {self.level}'")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self._name}: {self._level}'
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)
print(charlie_brown.level)
```
