---
id: 68c27fa1697ceb7c578b2c9f
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

Python에서 `@property` 데코레이터는 메서드를 속성으로 바꾸는 데 사용됩니다. 주로 속성 값을 가져오는 데 사용하는 getter 메서드를 정의할 때 사용합니다.

```py
class Person:
    def __init__(self, name):
        self._name = name

    @property
    def name(self):
        return self._name

p = Person('Alice')
print(p.name)  # Alice
```

`name`이라는 이름의 메서드를 `self` 매개변수와 함께 만들고 `@property`로 데코레이트하세요. 메서드 내부에서는 `self._name`를 반환하세요.
    
# --hints--

`Employee` 클래스에는 `name` 매개변수를 가진 `self` 메서드가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_args("self")`)) })
```

`name` 메서드는 `@property`로 데코레이트되어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_decorators("property")`)) })
```

`name` 메서드는 `self._name`를 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_return("self._name")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
