---
id: 68c2d2414f5e111c02f4db5e
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

`__repr__` 메서드는 `repr` 함수를 호출할 때 내부적으로 호출됩니다. 실제 동작을 확인하려면 코드 맨 아래에서 `repr(charlie_brown)` 호출 결과를 출력하세요.

# --hints--

`repr(charlie_brown)`를 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(repr(charlie_brown))")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level
--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)

--fcc-editable-region--
```
