---
id: 68c281cecda94ba2c6838e17
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

이제 `name`에 대한 getter를 정의했으니 `_name` 속성을 통해 `name` 속성에 접근할 수 있습니다. 그러니 `charlie_brown.name`을 콘솔에 출력하세요.

# --hints--

`charlie_brown.name`을 콘솔에 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown.name)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    @property
    def name(self):
        return self._name

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
