---
id: 68c27eb292bfec6b9402d58f
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

`__str__` 클래스에 `Employee` 메서드를 추가하세요. `name: level`와 `name`를 해당 속성으로 바꾸어 `level` 형식의 f-문자열을 반환하게 만드세요.

그다음 `charlie_brown`를 콘솔에 출력하세요.

# --hints--

`Employee` 클래스에는 `__str__` 매개변수를 가진 `self` 메서드가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_args("self")`)) })
```

`__str__` 메서드는 `name: level`와 `name` 값을 각각 `level`와 `self._name`로 바꾼 `self._level` 형식의 f-문자열을 반환해야 합니다.

```js
({ test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    empl = Employee("Frank", "dreamer")
    print(empl)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "Frank: dreamer" in output
`) })
```

`charlie_brown`를 콘솔에 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    

charlie_brown = Employee('Charlie Brown', 'trainee')

--fcc-editable-region--
```
