---
id: 68c7f7012a700243eff1cbc0
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

직원을 생성할 때 사용할 레벨은 특정 레벨 중에서 선택해야 합니다. `_base_salaries`를 사용해 레벨을 검증하고 직원의 적절한 급여를 설정하세요.

기존 `if` 문장 뒤에 `if`가 `level`에 없음을 확인하는 또 다른 `Employee._base_salaries`를 만드세요.

새로운 `if` 문장 안에서 `ValueError`를 `Invalid value '{level}' for 'level' attribute.` 메시지와 함께 발생시키세요. 여기서 `{level}`는 `level` 인수의 값으로 교체되어야 합니다.

# --hints--

`level`가 `Employee._base_salaries`에 없을 때 `ValueError`를 `Invalid value '{level}' for 'level' attribute.` 메시지와 함께 발생시켜야 하며, `{level}`는 `level` 인수의 값으로 교체되어야 합니다.

```js
({ test: () => runPython(`
  try:
    Employee("Frank", "dreamer")
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
    assert False, "Expected to raise ValueError with invalid level"
`) })
```

`level`가 `Employee._base_salaries`에 있을 때는 예외를 발생시키지 않아야 합니다.

```js
({ test: () => runPython(`
  levels = [
    "trainee",
    "junior",
    "mid-level",
    "senior"
  ]
  for level in levels:
    try:
      Employee("Frank", level)
    except Exception:
      assert False, f"Expected not to raise ValueError with level '{level}'"
    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }
--fcc-editable-region--
    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
