---
id: 68c7e19810faf3d00dba53d5
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

`_base_salaries`라는 이름의 클래스 속성을 만들고 다음 사전을 할당하세요:

```py
{
    'trainee': 1000,
    'junior': 2000,
    'mid-level': 3000,
    'senior': 4000
}
```

# --hints--

`Employee` 클래스에는 `_base_salaries`라는 이름의 클래스 속성이 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").has_variable("_base_salaries")`)) })
```

`{'trainee': 1000, 'junior': 2000, 'mid-level': 3000, 'senior': 4000}`를 `_base_salaries` 클래스 속성에 할당해야 합니다.

```js
({ test: () => runPython(`
  assert Employee._base_salaries == {'trainee': 1000, 'junior': 2000, 'mid-level': 3000, 'senior': 4000}
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    
--fcc-editable-region--
    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
