---
id: 68c80c1a8e58febb77120bbc
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

코드에서 마지막 `print` 호출을 제거하세요.

# --hints--

코드에서 `print(repr(charlie_brown))`을 제거해야 합니다.

```js
({ test: () => assert(runPython(`not _Node(_code).has_stmt("print(repr(charlie_brown))")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
print(repr(charlie_brown))
--fcc-editable-region--
```
