---
id: 68c8002beb157fe47d2e7699
title: 20단계
challengeType: 20
dashedName: step-20
---

# --description--

세터는 속성 값을 설정하는 데 사용하는 메서드로, 유효성 검사와 제한을 적용할 수 있습니다. `@propertyName.setter` 데코레이터를 사용해 세터를 만들 수 있으며, 여기서 `propertyName`은 설정할 속성 이름과 일치해야 합니다.

```py
class Person:
    def __init__(self, name):
        self._name = name

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, value):
        self._name = value

p = Person('Alice')
p.name = 'Abigail' # Calls the setter
print(p.name) # Abigail
```

게터 메서드 다음에 `name`와 `self` 매개변수를 가진 `new_name` 메서드를 만드세요. 이 메서드에 `@name.setter` 데코레이터를 붙이세요. 메서드 내부에서 `self._name`를 `new_name`로 설정하세요.

# --hints--

`Employee` 클래스에 `name`와 `self` 매개변수를 가진 `new_name` 메서드가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_args("self, new_name")`)) })
```

`name` 메서드는 `@name.setter` 데코레이터가 붙어 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_decorators("name.setter")`)) })
```

`name` 메서드는 `self._name`를 `new_name`로 설정해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_stmt("self._name = new_name")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    
--fcc-editable-region--
    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary


charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
