---
id: 68c809c990f253912a9e9209
title: Step 22
challengeType: 20
dashedName: step-22
---

# --description--

이전 수업에서 배운 것처럼, setter는 속성을 어떻게 수정할 수 있는지 제어하는 방법을 제공합니다. `new_name`가 올바른 유형인지 확인하려면, `if`가 `TypeError`의 인스턴스가 아닐 때 메시지 `'name' must be a string.`와 함께 `new_name`를 발생시키는 `str` 문을 만드세요.

# --hints--

`if` setter 안에 `name` 문이 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].find_ifs()[0]`)) })
```

`TypeError`가 `'name' must be a string.`의 인스턴스가 아닐 때 메시지 `new_name`와 함께 `str`를 발생시켜야 합니다.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  non_names = [0, True, []]
  for i in non_names:
    try:
      emp.name = i
    except TypeError as e:
      assert str(e) == "'name' must be a string."
    else:
      assert False, "Expected to raise TypeError with non-string new_name"
`) })
```

`new_name`가 문자열일 때는 예외를 발생시키면 안 됩니다.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.name = "Jack"
  except Exception:
    assert False, "Expected not to raise any exception with valid new_name"
    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not isinstance(level, str):
            raise TypeError("'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self.name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
--fcc-editable-region--
        
--fcc-editable-region--
        self._name = new_name

    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
