---
id: 68c9a1c241206043c81293d5
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

새로운 속성이 생겼으니, 이 속성에 대한 getter를 만드세요.

`salary` 매개변수를 가진 `self`라는 메서드를 만들고, 여기에 `@property` 데코레이터를 사용하세요. 메서드 내부에서는 `self._salary`를 반환하세요.

# --hints--

`Employee` 클래스에 `salary` 매개변수를 가진 `self` 메서드가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("salary").has_args("self")`)) })
```

`salary` 메서드는 `@property` 데코레이터가 붙어 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("salary").has_decorators("property")`)) })
```

`salary` 메서드는 `self._salary`를 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("salary").has_return("self._salary")`)) })
```


# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level
--fcc-editable-region--
    
--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
