---
id: 68c9cab4b1118da59eecfc56
title: 31단계
challengeType: 20
dashedName: step-31
---

# --description--

기존 `if` 문장 뒤에 `new_level`가 이미 선택된 레벨일 때 `ValueError`을 발생시키는 문장을 하나 더 만드세요.

초기화 시점에는 `_level`가 아직 존재하지 않으니, 비교하기 전에 `hasattr(self, '_level')`를 사용해 존재 여부를 확인하세요. 이렇게 하면 객체가 처음 생성될 때 `AttributeError`를 방지할 수 있습니다.

메시지에는 `'{level}' is already the selected level.`를 사용하고, `{level}`는 현재 레벨인 `{level}`로 교체하세요.

# --hints--

`level` 세터 안에 세 번째 `if` 문장이 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2]`)) })
```

`new_level`이 `self.level`과 같을 때, 메시지 `'{level}' is already the selected level.`와 함께 `ValueError`를 발생시켜야 하며, 여기서 `{level}`는 현재 레벨로 교체하세요.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.level = "trainee"
  except ValueError as e:
    assert str(e) == "'trainee' is already the selected level."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

`if` 문장은 비교하기 전에 `hasattr(self, '_level')`를 사용해 `_level`가 존재하는지 확인해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2].find_conditions()[0].is_equivalent("hasattr(self, '_level') and new_level == self.level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level
    

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
