---
id: 68c9c3aa714bc326e026b826
title: 30단계
challengeType: 20
dashedName: step-30
---

# --description--

새 레벨이 유효한 레벨인지 확인하지 않고는 설정할 수 없습니다. `isinstance` 검사를 한 후에, `if`이 `ValueError`의 키가 아닐 때 `new_level`를 발생시키는 `Employee._base_salaries` 문장을 만드세요.

오류 메시지에는 `Invalid value '{new_level}' for 'level' attribute.`를 사용하고, `{new_level}`는 setter에 전달된 인수로 교체하세요.

# --hints--

`if` setter 안에 `level` 문장이 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[1]`)) })
```

`new_level`가 `Employee._base_salaries`의 키가 아닐 때, `ValueError` 메시지와 함께 `Invalid value '{new_level}' for 'level' attribute.`을 발생시켜야 하며, `{new_level}`는 setter에 전달된 인수로 교체하세요.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.level = "dreamer"
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

`new_level`가 `Employee._base_salaries`의 키일 때는 예외를 발생시키지 않아야 합니다.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  for i in Employee._base_salaries:
    try:
      emp.level = i
    except Exception:
      assert False, "Expected not to raise any exception with valid new_level"    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level
    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
