---
id: 68c9a624fd54288f694ebdbf
title: 24단계
challengeType: 20
dashedName: step-24
---

# --description--

유효성 검사 과정을 확인하려면 `name` 속성에 문자열이 아닌 값을 자유롭게 설정하세요. 그 후에는 이름을 변경하는 코드를 모두 제거하여 `charlie_brown`의 `name`를 복원하세요.

# --hints--

`charlie_brown`의 `name`를 복원해야 합니다.

```js
({ test: () => runPython(`
  assert charlie_brown.name == 'Charlie Brown'
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
--fcc-editable-region--
charlie_brown.name = 'Snoopy'
--fcc-editable-region--
```
