---
id: 68c9a09163d4a42e9bc9b638
title: Step 24
challengeType: 20
dashedName: step-24
---

# --description--

이제 `_name` 속성에 대해 `name` 속성을 노출하는 setter 메서드가 있습니다.

그러니 `charlie_brown`의 `name`을 다른 이름으로 업데이트해 보세요. 일반 속성을 다루듯이 하면 됩니다. 이렇게 하면 내부적으로 setter 메서드가 호출됩니다.

# --hints--

`charlie_brown`의 `name`을 다른 이름으로 업데이트해야 합니다.

```js
({ test: () => runPython(`
  assert _Node(_code).has_variable("charlie_brown.name")
  assert charlie_brown.name != 'Charlie Brown'
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not isinstance(level, str):
            raise TypeError("'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self.name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
--fcc-editable-region--

--fcc-editable-region--
```
