---
id: 68c9a72c9dab42a15ce6972b
title: Step 26
challengeType: 20
dashedName: step-26
---

# --description--

이제 `_level` 속성에 대한 setter를 만드세요. `level`와 `self` 매개변수를 가진 `new_level`이라는 메서드를 만드세요. 이 메서드에 `@level.setter`를 데코레이터로 붙이세요.

메서드 내부에서 `self._level`를 `new_level`로 설정하세요.

# --hints--

`Employee` 클래스에 `level`와 `self` 매개변수를 가진 `new_level` 메서드가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].has_args("self, new_level")`)) })
```

`level` 메서드에 `@level.setter` 데코레이터가 붙어 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].has_decorators("level.setter")`)) })
```

`level` 메서드는 `self._level`를 `new_level`로 설정해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].has_stmt("self._level = new_level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not isinstance(level, str):
            raise TypeError("'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self.name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

--fcc-editable-region--
    
    
--fcc-editable-region--

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
