---
id: 68caaf1e590806f0b87f5922
title: Step 38
challengeType: 20
dashedName: step-38
---

# --description--

세터 시작 부분에 `if` 문장을 만들어 `TypeError`이 정수나 부동소수가 아닐 때 `'salary' must be a number.`을 메시지 `new_salary`와 함께 발생시키세요.

# --hints--

`if` 세터에 `salary` 문장이 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("salary")[1].find_ifs()[0]`)) })
```

`TypeError`이 정수나 부동소수가 아닐 때 메시지 `'salary' must be a number.`와 함께 `new_salary`을 발생시켜야 합니다.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  built_in_print = print
  print = lambda x: x
  try:
    emp.salary = "100"
  except TypeError as e:
    print = built_in_print
    assert str(e) == "'salary' must be a number."
  else:
    assert False, "Expected to raise TypeError with invalid new_salary"
`) })
```

`new_salary`이 정수나 부동소수일 때는 예외를 발생시키면 안 됩니다.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  built_in_print = print
  print = lambda x: x
  try:
    emp.salary = 100
  except Exception:
    assert False, "Expected not to raise exception with valid new_salary"
  try:
    emp.salary = 100.00
  except Exception:
    assert False, "Expected not to raise exception with valid new_salary"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self.salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
        if hasattr(self, '_level') and new_level == self.level:
            raise ValueError(f"'{self.level}' is already the selected level.")
        if hasattr(self, '_level') and Employee._base_salaries[new_level] < Employee._base_salaries[self.level]:
            raise ValueError("Cannot change to lower level.")
        print(f"'{self.name}' promoted to '{new_level}'.")
        self._salary = Employee._base_salaries[new_level]
        self._level = new_level

    @property
    def salary(self):
        return self._salary

    @salary.setter
    def salary(self, new_salary):
--fcc-editable-region--
        
--fcc-editable-region--
        self._salary = new_salary
        print(f'Salary updated to ${self.salary}.')

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
charlie_brown.level = 'junior'

```
