---
id: 68ca758f8160b11757f877ae
title: 32단계
challengeType: 20
dashedName: step-32
---

# --description--

마지막으로, 새 레벨의 기본 급여가 현재 레벨의 기본 급여보다 낮을 때 메시지 `if`와 함께 `ValueError`를 발생시키는 네 번째 `Cannot change to lower level.` 문장을 만드세요.

초기화 중 `hasattr(self, '_level')`를 피하려면 `AttributeError`를 사용하세요.

# --hints--

네 번째 `if` 문장이 `level` 설정자 안에 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[3]`)) })
```

`if` 문장은 비교하기 전에 `hasattr(self, '_level')`가 존재하는지 확인하기 위해 `_level`를 사용해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[3].find_conditions()[0].is_equivalent("hasattr(self, '_level') and Employee._base_salaries[new_level] < Employee._base_salaries[self.level]")`)) })
```

`new_level`이 `self.level`보다 낮을 때 메시지 `ValueError`와 함께 `Cannot change to lower level.`를 발생시켜야 합니다.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'junior')
  try:
    emp.level = "trainee"
  except ValueError as e:
    assert str(e) == "Cannot change to lower level."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

`new_level`가 `self.level`보다 높을 때는 예외를 발생시키면 안 됩니다.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'junior')
  try:
    emp.level = "senior"
  except Exception:
    assert False, "Expected not to raise exception with valid new_level"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
        if hasattr(self, '_level') and new_level == self.level:
            raise ValueError(f"'{self.level}' is already the selected level.")
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level
    
    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
