---
id: 68ca758f8160b11757f877ae
title: 28단계
challengeType: 20
dashedName: step-28
---

# --description--

마지막으로, 새 레벨의 기본 급여가 현재 레벨의 기본 급여보다 낮을 때 `if` 메시지와 함께 `ValueError`를 발생시키는 세 번째 `Cannot change to lower level.` 문을 만드세요.

# --hints--

세 번째 `if` 문이 `level` 설정자 안에 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2]`)) })
```

`new_level`가 `self.level`보다 낮을 때 `ValueError` 메시지와 함께 `Cannot change to lower level.`를 발생시켜야 합니다.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'junior')
  try:
    emp.level = "trainee"
  except ValueError as e:
    assert str(e) == "Cannot change to lower level."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

`new_level`가 `self.level`보다 높을 때는 예외를 발생시키면 안 됩니다.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'junior')
  try:
    emp.level = "senior"
  except Exception:
    assert False, "Expected not to raise exception with valid new_level"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level
--fcc-editable-region--
    @level.setter
    def level(self, new_level):
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
        if new_level == self.level:
            raise ValueError(f"'{self.level}' is already the selected level.")
        
        self._level = new_level
    
--fcc-editable-region--
    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
