---
id: 68caacb0f4311cc9be9a2132
title: Step 36
challengeType: 20
dashedName: step-36
---

# --description--

이제 `salary` 세터 코딩에 집중하세요. `salary` 게터 다음에 `salary` 속성에 대한 간단한 세터를 만들어 메서드에 인수로 전달된 값을 `self._salary`에 설정하세요. 그다음 `Salary updated to $`를 출력한 뒤 새 급여와 마침표를 출력하세요.

새 급여 검증은 다음 단계에서 처리할 것입니다.

# --hints--

`Employee` 클래스에는 `salary`로 장식된 `@salary.setter` 메서드가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("salary")[1].has_decorators("salary.setter")`)) })
```

`salary` 세터는 인수로 전달된 값을 `self._salary`에 설정해야 합니다.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  built_in_print = print
  print = lambda x: x
  emp.salary += 100
  print = built_in_print
  assert emp.salary == Employee._base_salaries['trainee'] + 100
`) })
```

`salary` 세터는 새 급여와 마침표를 출력하기 전에 `Salary updated to $`를 출력해야 합니다.

```js
({ test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    emp = Employee('Frank', 'trainee')
    emp.salary += 100
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert f"Salary updated to \${Employee._base_salaries['trainee'] + 100}" in output
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
        if hasattr(self, '_level') and new_level == self.level:
            raise ValueError(f"'{self.level}' is already the selected level.")
        if hasattr(self, '_level') and Employee._base_salaries[new_level] < Employee._base_salaries[self.level]:
            raise ValueError("Cannot change to lower level.")
        print(f"'{self.name}' promoted to '{new_level}'.")
        self._salary = Employee._base_salaries[new_level]
        self._level = new_level

    @property
    def salary(self):
        return self._salary

--fcc-editable-region--
    
--fcc-editable-region--

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')

charlie_brown.level = 'junior'

```
