---
id: 69b014adbc624861e6ff3dd1
title: 21단계
challengeType: 20
dashedName: step-21
---

# --description--

이제 `name` setter가 만들어졌으니, `__init__`을 `self.name = name` 대신 `self._name = name`를 사용하도록 업데이트하세요. 이렇게 하면 `name` 할당이 setter를 통해 이루어지고, 다음 단계에서 유효성 검사가 추가될 것입니다.

또한 `if` 조건을 `level`만 검사하도록 조정하고, `TypeError` 메시지도 그에 맞게 업데이트하세요.

# --hints--

`__init__` 메서드는 `self.name = name` 대신 `self._name = name`를 사용해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

`__init__` 메서드는 `isinstance(level, str)`만 검사하고 `TypeError` 메시지를 그에 맞게 업데이트해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_if("not isinstance(level, str)").find_body().has_stmt("raise TypeError(\\"'level' attribute must be of type 'str'.\\")")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }
    def __init__(self, name, level):
--fcc-editable-region--
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]
--fcc-editable-region--

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        self._name = new_name
        
    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary


charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
