---
id: 69b023541591abde576613db
title: 27단계
challengeType: 20
dashedName: step-27
---

# --description--

이제 `level` 세터가 만들어졌으니, `__init__`을 `self.level = level` 대신 `self._level = level`를 사용하도록 업데이트하세요. 이렇게 하면 `level` 할당이 세터를 통해 이루어지고, 다음 단계에서 유효성 검사가 추가될 것입니다.

또한, 세터가 다음 단계에서 유효성 검사를 처리할 것이므로 `if not isinstance(level, str)`에서 `__init__` 검사를 제거하세요.

# --hints--

`__init__` 메서드는 `self.level = level` 대신 `self._level = level`를 사용해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.level = level")`)) })
```

`__init__` 메서드에는 `isinstance(level, str)` 검사가 포함되어 있으면 안 됩니다.

```js
({ test: () => assert(runPython(`not _Node(_code).find_class("Employee").find_function("__init__").find_if("not isinstance(level, str)").find_body().has_stmt("raise TypeError(\\"'level' attribute must be of type 'str'.\\")")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
--fcc-editable-region--
        if not isinstance(level, str):
            raise TypeError("'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self.name = name
        self._level = level
        self._salary = Employee._base_salaries[level]
--fcc-editable-region--

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        self._level = new_level

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
