---
id: 60eebd07ea685b0e777b5583
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Welcome to the Registration Form project! 首先，在文檔的頂部添加 `!DOCTYPE html` 聲明，這樣瀏覽器就知道它在解析什麼類型的文檔。

# --hints--

代碼應包含 `DOCTYPE` 聲明。

```js
assert(code.match(/<!DOCTYPE/gi));
```

應該在 `DOCTYPE` 聲明之後包含一個空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

應該將文檔類型定義爲 `html`。

```js
assert(code.match(/html/gi));
```

你應該在類型之後使用 `>` 關閉 `DOCTYPE` 聲明。

```js
assert(code.match(/html\s*>/gi));
```

`DOCTYPE` 聲明應位於 HTML 的開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
