---
id: 60f02e7361b68405e27b62a5
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

在 `body` 中，為內容提供標題背景關係，方法是添加一個帶有文字 `Registration Form` 的 `h1`。

# --hints--

你應該在 `body` 中添加 `h1`。

```js
assert.exists(document.querySelector('body > h1'));
```

你應該給 `h1` 設定文字 `Registration Form`。

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```
