---
id: 60f0286404aefb0562a4fdf9
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

在 `head` 元素內添加 `title` 和 `meta` 元素。為你的專案設定標題為 `Registration Form`，並將 `charset` 屬性設為 `utf-8`，添加到你的 `meta` 元素中。

# --hints--

你的程式碼應該有一個 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

你的 `title` 元素應該位於 `head` 元素內。

```js
assert.exists(document.querySelector('head > title'));
```

你的專案應該有一個標題為 `Registration Form`。請記住，標題的大小寫和拼寫很重要。

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

你應該在 `head` 元素內創建一個 `meta` 元素。

```js
assert.exists(document.querySelector('head > meta'));
```

你應該為 `meta` 元素指定一個 `charset` 屬性，值為 `utf-8`。

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

你的 `meta` 元素應該是空元素，它沒有結束標籤 `</meta>`。

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
