---
id: 60f027099a15b00485563dd2
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `DOCTYPE` 之下，添加一個帶有 `lang` 屬性且設為 `en` 的 `html` 元素，這樣你就有一個開始放置程式碼的地方。

# --hints--

你的 `DOCTYPE` 宣告應該放在你的 HTML 開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你的 `html` 元素應該有一個開啟標記，以及 `lang` 屬性，其值為 `en`。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素應該有一個結束標籤。

```js
assert(code.match(/<\/html\s*>/));
```

你的 `html` 頁籤應該按照正確的順序排列。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

你應該只擁有一個 `html` 元素。

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
