---
id: 60f1a5e2d2c23707a4f9a660
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

現在，透過將某些瀏覽器為 `body` 預設添加的 `margin` 設為 `0`，來移除水平捲軸。

# --hints--

你應該在 `body` 元素選擇器內添加 `margin`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.margin);
```

你應該給 `margin` 一個 `0` 的值。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
}
--fcc-editable-region--
```
