---
id: 60f5d2776c854e069560fbe6
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

第一個 `fieldset` 將包含姓名、電子郵件和密碼欄。首先為第一個 `fieldset` 添加四個 `label` 元素。

# --hints--

總共應該有三個 `fieldset` 元素。

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

所有 `fieldset` 元素都應該是 `form` 元素的直接子元素。

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

四個 `label` 元素都應該是第一個 `fieldset` 元素的直接子元素。請確保你關閉 `label` 元素。

```js
assert.equal(document.querySelectorAll('form fieldset:first-child > label')?.length, 4);
```

最後兩個 `fieldset` 元素應該是空的。

```js
const LastTwoFieldsetElements = Array.from(document.querySelectorAll("form fieldset:not(fieldset:first-child)"));
assert.isTrue(LastTwoFieldsetElements.filter((ele) => ele.innerHTML?.replace(/\s/g, ""))?.length === 0);
```

請確定你關閉 `label` 元素。

```js
assert.lengthOf(document.querySelector('fieldset')?.innerHTML?.match(/<\/label>/g), 4)
```



# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
--fcc-editable-region--
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>

      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
