---
id: 60fabf0dd4959805dbae09e6
title: 步驟 33
challengeType: 0
dashedName: step-33
---

# --description--

在第一個 `label` 之後添加另一個 `label`，文字為 `Input your age (years): `。然後，巢狀一個 `input`，其 `type` 為 `number`。

接著，為 `input` 添加一個值為 `13` 的 `min` 屬性，因為 13 歲以下的使用者不應該註冊。此外，使用者可能不會超過 120 歲；為 `input` 添加一個值為 `120` 的 `max` 屬性。

現在，如果有人嘗試提交帶有超出範圍值的表單，將會出現警告，且表單不會提交。試試看。

# --hints--

你應該在第三個 `fieldset` 的現有 `label` 之後添加一個 `label`。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label + label'));
```

你應該給 `label` 指定文字 `Input your age (years): `。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2)')?.textContent?.trim(), 'Input your age (years):');
```

你應該給 `label` 一個 `type` 為 `number` 的 `input`。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]'));
```

你應該為 `input` 給予一個值為 `13` 的 `min` 屬性。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.min, '13');
```

你應該為 `input` 給予一個值為 `120` 的 `max` 屬性。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.max, '120');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label>Upload a profile picture: <input type="file" /></label>
        
      </fieldset>
--fcc-editable-region--
      <label for="terms-and-conditions">
        <input id="terms-and-conditions" type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
      </label>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
