---
id: 60fad1cafcde010995e15306
title: 步驟 42
challengeType: 0
dashedName: step-42
---

# --description--

在表單提交時，為每個可提交的元素提供 `name` 屬性是有用且良好的做法。此屬性用於識別表單提交中的元素。

除了兩個已經命名的 `radio` 輸入外，為每個可提交的元素指定你選擇的唯一 `name` 屬性。

# --hints--

你應該給期望名字的 `input` 一個 `name` 屬性。*附註：我本來會選擇 `first-name`*

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(1) > input')?.name);
```

你應該給期望姓氏的 `input` 添加一個 `name` 屬性。*附註：我本來會選擇 `last-name`*

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(2) > input')?.name);
```

你應該為 `email` 添加 `name` 屬性。*PS 我本來會選擇 `email`*

```js
assert.isNotEmpty(document.querySelector('input[type="email"]')?.name);
```

你應該給 `password` 一個 `name` 屬性。*PS 我本來會選擇 `password`*

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.name);
```

你應該給 `checkbox` 一個 `name` 屬性。*PS 我本來會選擇 `terms`*

```js
assert.isNotEmpty(document.querySelector('input[type="checkbox"]')?.name);
```

你應該給 `file` 一個 `name` 屬性。*PS 我本來會選擇 `file`*

```js
assert.isNotEmpty(document.querySelector('input[type="file"]')?.name);
```

你應該給 `number` 一個 `name` 屬性。*PS 我本來會選擇 `age`*

```js
assert.isNotEmpty(document.querySelector('input[type="number"]')?.name);
```

你應該給 `select` 元素一個 `name` 屬性。*PS 我本來會選擇 `referrer`*

```js
assert.isNotEmpty(document.querySelector('select')?.name);
```

你應該為 `textarea` 元素指定一個 `name` 屬性。*PS 我本來會選擇 `bio`*

```js
assert.isNotEmpty(document.querySelector('textarea')?.name);
```

你不應該給任何 `option` 元素添加 `name` 屬性。

```js
[...document.querySelectorAll('option')]?.forEach(option => assert.isUndefined(option?.name));
```

你不應該給任何 `label` 元素添加 `name` 屬性。

```js
[...document.querySelectorAll('label')]?.forEach(label => assert.isUndefined(label?.name));
```

你不應該給任何 `fieldset` 元素添加 `name` 屬性。

```js
[...document.querySelectorAll('fieldset')]?.forEach(fieldset => assert.isEmpty(fieldset?.name));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
        <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset>
        <label for="profile-picture">Upload a profile picture: <input id="profile-picture" type="file"/></label>
        <label for="age">Input your age (years): <input id="age" type="number" min="13" max="120" /></label>
        <label for="referrer">How did you hear about us?
          <select id="referrer">
            <option value="">(select one)</option>
            <option value="1">freeCodeCamp News</option>
            <option value="2">freeCodeCamp YouTube Channel</option>
            <option value="3">freeCodeCamp Forum</option>
            <option value="4">Other</option>
          </select>
        </label>
        <label for="bio">Provide a bio:
          <textarea id ="bio" rows="3" cols="30" placeholder="I like coding on the beach..."></textarea>
        </label>
      </fieldset>
      <label for="terms-and-conditions">
        <input id="terms-and-conditions" type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
      </label>
      --fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
