---
id: 60fab8367d35de04e5cb7929
title: 步驟 31
challengeType: 0
dashedName: step-31
---

# --description--

在新添加的 `label` 中，緊接著 `input` 元素添加文字 `I accept the terms and conditions`。然後將文字 `terms and conditions` 連結到以下位置：

```md
https://www.freecodecamp.org/news/terms-of-service/
```

# --hints--

你應該在第三個 `fieldset` 之後的 `label` 中添加 `I accept the terms and conditions` 文字。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label')?.innerText.trim(), 'I accept the terms and conditions');
```

你應該使用 `a` 元素來聯結條款和條件。 

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label a'));
```

你應該將新文字緊接在 `label` 中的 `input` 元素之後。 

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input + a'));
```

你應該給 `a` 元素一個 `href`，其值為 `https://www.freecodecamp.org/news/terms-of-service/`。

```js
assert.match(document.querySelector('fieldset:nth-child(3) + label > input + a')?.href, /https:\/\/www\.freecodecamp\.org\/news\/terms-of-service\/?/);
```

你應該只將 `a` 元素包裹在文字 `terms and conditions` 周圍。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input + a')?.textContent.trim(), 'terms and conditions');
```

你的錨點元素內的文字有多餘的前置或後置空白。錨點文字中唯一的空格應該是在字、單字 `terms`、`and` 和 `conditions` 之間。

```js
const nestedAnchor = document.querySelector('fieldset:nth-child(3) + label > input + a')?.textContent;
const innerContent = nestedAnchor?.innerHTML;
assert.isNotTrue(/^\s+|\s+$/.test(innerContent));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--
      <label for="terms-and-conditions"><input id="terms-and-conditions" type="checkbox" required /></label>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
