---
id: 62cc5b1779e4d313466f73c5
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

將 `link` 元素巢狀於 `head` 元素內。為它添加一個值為 `stylesheet` 的 `rel` 屬性和一個值為 `styles.css` 的 `href` 屬性。

# --hints--

你的程式碼應該有一個 `link` 元素。

```js
assert.isNotNull(document.querySelector('link'));
```

你不應該更改你現有的 `head` 頁籤。請確保你沒有刪除結束標籤。

```js
const heads = document.querySelectorAll('head');
assert.strictEqual(heads?.length, 1);
```

你應該有一個 `link` 元素。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

你的 `link` 元素應該是空元素，不應該有結束標籤 `</link>`。

```js
assert.notMatch(code, /<\/link>/);
```

你的 `link` 元素應該在你的 `head` 元素內。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

你的 `link` 元素應該有一個 `rel` 屬性，其值為 `stylesheet`。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

你的 `link` 元素應該有一個 `href` 屬性，其值為 `styles.css`。

```js
const link = document.querySelector('link');
assert.strictEqual(link?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
