---
id: 62ff8b9dab5ac88e4d3d43a3
title: 步驟 18
challengeType: 0
dashedName: step-18
---

# --description--

遵循無障礙最佳實務，使用 `for` 屬性將 `input` 元素和 `label` 元素連結在一起。

將 `first-name`、`last-name`、`email` 和 `new-password` 用作相應 `id` 屬性的值。

# --hints--

第一個 `input` 元素應該有一個 `id` 為 `first-name`。

```js
assert(document.querySelectorAll('input')?.[0]?.matches('#first-name'))
```

第二個 `input` 元素應該有一個 `id` 為 `last-name`。

```js
assert(document.querySelectorAll('input')?.[1]?.matches('#last-name'))
```

第三個 `input` 元素應該有一個 `id` 為 `email`。

```js
assert(document.querySelectorAll('input')?.[2]?.matches('#email'))
```

第四個 `input` 元素應該有一個 `id` 為 `new-password`。

```js
assert(document.querySelectorAll('input')?.[3]?.matches('#new-password'))
```

第一個 `label` 元素應該有一個 `for` 屬性，值為 `first-name`。

```js
assert(document.querySelectorAll('label')?.[0]?.matches('label[for="first-name"]'))
```

第二個 `label` 元素應該有一個 `for` 屬性，值為 `last-name`。

```js
assert(document.querySelectorAll('label')?.[1]?.matches('label[for="last-name"]'))
```

第三個 `label` 元素應該有一個 `for` 屬性，值為 `email`。

```js
assert(document.querySelectorAll('label')?.[2]?.matches('label[for="email"]'))
```

第四個 `label` 元素應該有一個 `for` 屬性，值為 `new-password`。

```js
assert(document.querySelectorAll('label')?.[3]?.matches('label[for="new-password"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name: <input /></label>
        <label>Enter Your Last Name: <input /></label>
        <label>Enter Your Email: <input /></label>
        <label>Create a New Password: <input /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
