---
id: 62ff8e998d3e7eae14d6ae3b
title: 第 30 步
challengeType: 0
dashedName: step-30
---

# --description--

你需要確認使用者已閱讀條款和條件。

添加一個 `label` 元素。在新創建的 `label` 元素內添加一個 `input` 元素，並將 `type` 屬性設為 `checkbox`。將此 `input` 元素設為 `required`，以便使用者在未同意條款和條件時無法註冊。

為元素添加 `id` 和 `for` 屬性，值為 `terms-and-conditions`，以提升存取性。

# --hints--

你應該在第三個 `fieldset` 元素之後添加一個 `label`。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label'));
```

你應該為你的 `label` 元素添加一個 `input`。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input'));
```

你應該為 `input` 元素添加一個值為 `checkbox` 的 `type` 屬性。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.type, 'checkbox');
```

你應該為你的 `input` 元素添加 `required` 屬性。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.required, true);
```

`input` 元素應該有一個 `id` 為 `terms-and-conditions`。

```js
assert(document.querySelector('fieldset:nth-child(3) + label > input')?.matches('#terms-and-conditions'))
```

`label` 元素應該有一個 `for` 屬性，其值為 `terms-and-conditions`。

```js
assert(document.querySelector('fieldset:nth-child(3) + label')?.matches('label[for="terms-and-conditions"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--

--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
