---
id: 63507c4b63731437227b0134
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Die `.map()`-Methode nimmt eine Callback-Funktion als erstes Argument. Diese Callback-Funktion erhält einige Argumente, aber das erste ist das aktuell verarbeitete Element. Hier ist ein Beispiel:

```js
array.map(el => {

})
```

Die Callback-Funktion muss einen Wert zurückgeben. In diesem Fall möchten Sie den Wert jedes Elements in eine Zahl umwandeln und zurückgeben. Das können Sie erreichen, indem Sie den `Number()`-Konstruktor verwenden und das Element als Argument übergeben.

Fügen Sie Ihrer `.map()`-Methode eine Callback-Funktion hinzu, die jedes Element in eine Zahl umwandelt.

# --hints--

Ihre `.map()`-Methode sollte eine Callback-Funktion enthalten.

```js
assert.match(calculate.toString(), /array\.map\(\s*(\(\s*\w+\s*\)|\w+)\s*=>|array\.map\(\s*function\s*\(\s*\w+\s*\)\s*\{|array\.map\(\s*Number\s*\)/)
```

Ihre Callback-Funktion sollte den `Number`-Konstruktor verwenden, um jedes Array-Element in eine Zahl umzuwandeln.

```js
assert.match(calculate.toString(), /Number\(\s*\w+\s*\)|Number/);
```

Ihre Callback-Funktion sollte das Schlüsselwort `new` nicht mit dem `Number`-Konstruktor verwenden.

```js
assert.notMatch(calculate.toString(), /new/);
```

Ihre Callback-Funktion sollte das Element in eine Zahl umgewandelt zurückgeben.

```js
assert.match(calculate.toString(), /(array\.map\(\s*(\(\s*\w+\s*\)|\w+)\s*=>|array\.map\(\s*function\s*\(\s*\w+\s*\)\s*\{)\s*(return\s+)?Number\(\s*\w+\s*\)|array\.map\(\s*Number\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>

    <form onsubmit="calculate();">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
--fcc-editable-region--
const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map();
}
--fcc-editable-region--
```
