---
id: 6352ec8b9c70fd17b8c7ba3f
title: Schritt 40
challengeType: 0
dashedName: step-40
---

# --description--

Jetzt müssen Sie die Werte richtig sortieren. Verketten Sie die `.sort()`-Methode an Ihren `Object.keys()`-Aufruf.

Für den Callback müssen Sie das `counts`-Objekt verwenden, um die Werte jedes Schlüssels zu vergleichen. Sie können die `a`- und `b`-Parameter verwenden, um auf die Schlüssel zuzugreifen. Geben Sie dann den Wert von `counts[b]` minus den Wert von `counts[a]` zurück.

Greifen Sie abschließend mit der Klammernotation auf das erste Element im Array zu, um Ihre `highest`-Variable zu vervollständigen.

# --hints--

Ihre `getMode`-Funktion sollte die `sort`-Methode verwenden, um das `Object.keys()`-Array zu sortieren.

```js
assert.match(getMode.toString(), /highest\s*=\s*Object\.keys\(\s*counts\s*\)\.sort\(/)
```

Ihre `getMode`-Funktion sollte einen Callback an die `sort`-Methode mit den Parametern `a` und `b` übergeben.

```js
const sortCallbackRegex = __helpers.functionRegex(null, ['a', 'b'], {
  includeBody: false
});
assert.match(
  getMode.toString(),
  __helpers.concatRegex(
    /highest\s*=\s*Object\.keys\(\s*counts\s*\)\.sort\(\s*/,
    sortCallbackRegex
  )
);
```

Ihre `getMode`-Funktion sollte die `sort`-Methode verwenden, um das `Object.keys()`-Array zu sortieren und den Wert von `counts[b]` minus `counts[a]` zurückgeben.

```js
assert.match(
  getMode.toString(),
  __helpers.concatRegex(
    /highest\s*=\s*Object\.keys\(\s*counts\s*\)\.sort\(\s*/,
    __helpers.functionRegex(null, ['a', 'b'], { includeBody: false }),
    /\s*(?:return\s+)?counts\s*\[\s*b\s*\]\s*-\s*counts\s*\[\s*a\s*\]/
  )
);
```

Ihre `highest`-Variable sollte den Wert des ersten Eintrags im sortierten `Object.keys(counts)`-Array haben.

```js
assert.match(getMode.toString(), /highest\s*=\s*Object\.keys\(\s*counts\s*\)\.sort\([\s\S]*\)\s*\[\s*0\s*\]\s*;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>
    <form onsubmit="calculate(); return false;">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
const getMean = (array) => array.reduce((acc, el) => acc + el, 0) / array.length;

const getMedian = (array) => {
  const sorted = array.toSorted((a, b) => a - b);
  const median =
    sorted.length % 2 === 0
      ? getMean([sorted[sorted.length / 2], sorted[sorted.length / 2 - 1]])
      : sorted[Math.floor(sorted.length / 2)];
  return median;
}

--fcc-editable-region--
const getMode = (array) => {
  const counts = {};
  array.forEach(el => counts[el] = counts[el] ? counts[el] + 1 : 1);

  if (new Set(Object.values(counts)).size === 1) {
    return null;
  }
  const highest = Object.keys(counts);
}
--fcc-editable-region--


const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el)).filter(el => !isNaN(el));
  
  const mean = getMean(numbers);
  const median = getMedian(numbers);

  document.querySelector("#mean").textContent = mean;
  document.querySelector("#median").textContent = median;
}
```
