---
id: 661899113600ef307d8a7487
title: Schritt 23
challengeType: 0
dashedName: step-23
---

# --description--

Um den Median eines Arrays mit einer ungeraden Anzahl von Elementen zu erhalten, müssen Sie die mittlere Zahl finden und zurückgeben.

So finden Sie die mittlere Zahl eines Arrays mit einer ungeraden Anzahl von Elementen:

```js
arr[Math.floor(arr.length / 2)];
```

Hier ist ein längeres Beispiel, das die mittlere Zahl eines Arrays mit `5` Elementen findet:

```js
const numbers = [1, 2, 3, 4, 5];
const middleNumber = numbers[Math.floor(numbers.length / 2)];
console.log(middleNumber); // 3
```

Der Grund, warum Sie `Math.floor` verwenden, ist, dass das Ergebnis der Division einer ungeraden Zahl durch `2` eine Dezimalzahl sein wird. `Math.floor` rundet auf die nächstkleinere ganze Zahl ab.

Deklarieren Sie eine `oddListMedian`-Variable und weisen Sie ihr das Ergebnis der Suche nach der mittleren Zahl des `testArr1` zu. Geben Sie dann die `oddListMedian`-Variable in der Konsole aus.

Öffnen Sie die Konsole, um das Ergebnis zu sehen.

# --hints--

Sie sollten eine `oddListMedian`-Variable haben.

```js
assert.isDefined(oddListMedian);
```

Sie sollten das Ergebnis der Suche nach der mittleren Zahl für das `testArr1` der `oddListMedian`-Variable zuweisen. Denken Sie daran, sich auf das frühere Beispiel zu beziehen: `arr[Math.floor(arr.length / 2)];`

```js
assert.strictEqual(oddListMedian, 3);
```

Sie sollten eine Konsolenausgabe haben, die die `oddListMedian`-Variable ausgibt.

```js
assert.match(code, /console\.log\(\s*oddListMedian\s*\);?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>
    <form onsubmit="calculate(); return false;">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
const getMean = (array) => array.reduce((acc, el) => acc + el, 0) / array.length;

--fcc-editable-region--
const testArr1 = [1, 2, 3, 4, 5];
const testArr2 = [1, 2, 3, 4, 5, 6];
const isEven = testArr2.length % 2 === 0;
console.log(isEven);


--fcc-editable-region--
const getMedian = (array) => {
  const sorted = array.sort((a, b) => a - b);
}

const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el)).filter(el => !isNaN(el));
  
  const mean = getMean(numbers);

  document.querySelector("#mean").textContent = mean;
}
```
