---
id: 5f331e55dfab7a896e53c3a1
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

The `title` is one of several elements that provide extra information not visible on the web page, but it is useful for search engines or how the page gets displayed.

Dentro do elemento `head`, adicione um elemento `meta` com um atributo chamado `charset` definido para o valor `utf-8` para dizer ao navegador como codificar caracteres para a página.

Remember that `meta` elements are void elements that require no ending tag.

# --hints--

Você deve criar uma tag `meta`.

```js
assert.match(code, /<meta.*?\/?>/is);
```

O elemento `meta` deve ser um elemento nulo, sem uma tag final `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

A tag `meta` deve ter um atributo `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Your `charset` attribute should have a value of `utf-8`.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
