---
id: 5f33071498eb2472b87ddee4
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this project, you will learn the basics of CSS (Cascading Style Sheets) by building a cafe menu. CSS is the language used to style an HTML document. It describes how HTML elements should be displayed on the screen.

As you learned in the last few steps of the Cat Photo App, there is a basic structure needed to start building your web page. Cada documento HTML deve ter uma declaração `DOCTYPE` e um elemento `html`. O `DOCTYPE` diz ao navegador em que versão do HTML o documento está. O elemento `html` representa o elemento raiz que contém todos os outros elementos.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

Add the `<!DOCTYPE html>` tag, and an `html` element with a `lang` attribute of `en`.

# --hints--

Você deve ter uma declaração `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Você deve ter uma tag de abertura `<html lang="en">` com o idioma definido para inglês.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Você deve acrescentar uma tag de fechamento `</html>`.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
