---
id: 6852f01116912861c7d1b0e6
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Dentro do método `__init__`, você precisa atribuir os valores dos parâmetros aos atributos do objeto para que cada email possa armazenar suas informações.

Dentro do seu método `__init__`, remova a palavra-chave `pass` e atribua os parâmetros a `self.sender`, `self.receiver`, `self.subject` e `self.body`.

# --hints--

Você deve atribuir o parâmetro `sender` a `self.sender`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.sender").is_equivalent("self.sender = sender")`))
    }

})
```

Você deve atribuir o parâmetro `receiver` a `self.receiver`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.receiver").is_equivalent("self.receiver = receiver")`))
    }

})
```

Você deve atribuir o parâmetro `subject` a `self.subject`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.subject").is_equivalent("self.subject = subject")`))
    }

})
```

Você deve atribuir o parâmetro `body` a `self.body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.body").is_equivalent("self.body = body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        pass
--fcc-editable-region--
```
