---
id: 6852f11d412c29665985a915
title: Passo 12
challengeType: 20
dashedName: step-12
---

# --description--

Agora que os usuários podem enviar emails, eles precisam de uma forma de armazenar os emails que recebem. Cada usuário deve ter uma caixa de entrada para coletar seus emails.

Adicione um atributo `inbox` ao método `__init__` da classe `User` e inicialize-o como uma lista vazia `[]`. Isso armazenará todos os e-mails que o usuário receber.

# --hints--

Você deve adicionar `self.inbox = []` ao método `__init__` da classe `User`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
```
