---
id: 6852ea0c3cadb54e897acdac
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Neste workshop, você vai construir um *Email Simulator* que simula o envio, recebimento e gerenciamento de emails entre diferentes usuários. Você vai aprender sobre classes, objetos e como organizar o código de forma orientada a objetos.

Comece criando uma classe chamada `Email` usando a palavra-chave `class`.

# --hints--

Você deve criar uma classe chamada `Email`.

```js
({
    test: () => 
    {
        assert(runPython(`(_Node(_code).has_class("Email"))`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
