---
id: 6852ef95ccc27660275692cd
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Sua classe `Email` precisa armazenar informações sobre cada email quando ele é criado.

Dentro da sua classe `Email`, remova a palavra-chave `pass` existente e crie o método `__init__` que recebe `self`, `sender`, `receiver`, `subject` e `body` como parâmetros.

# --hints--

Você deve definir um método `__init__` dentro da classe `Email`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__init__")`))
    }

})
```

Seu método `__init__` deve ter os parâmetros: `self`, `sender`, `receiver`, `subject` e `body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").has_args("self, sender, receiver, subject, body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    pass
--fcc-editable-region--
```
