---
id: 6852f0c3fcedb964ab959921
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Agora vamos testar se o atributo `read` foi adicionado corretamente à sua classe `Email`. Como você já tem um objeto de email dos passos anteriores, imprima o atributo `read` para ver que ele agora é `False` por padrão.

# --hints--

Você deve imprimir `email_obj.read` para testar o atributo read.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)

--fcc-editable-region--
```
