---
id: 6852f06fdde1b4634a263f03
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Agora vamos testar a classe `Email` criando um objeto de email e verificando se ele armazena as informações corretamente. Você vai imprimir alguns atributos para verificar se tudo funciona.

Crie um objeto de email chamado `email_obj` com `alice@example.com` como remetente, `bob@example.com` como destinatário, `Hello` como assunto e `Hi Bob!` como corpo. Em seguida, imprima os atributos `sender` e `subject` em declarações de impressão separadas para verificar se estão armazenados corretamente.

# --hints--

Você deve criar um objeto de email chamado `email_obj` com os parâmetros especificados.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_variable("email_obj").is_equivalent("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`))
    }

})
```

Você deve imprimir `email_obj.sender`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.sender)")`))
    }

})
```

Você deve imprimir `email_obj.subject`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.subject)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--

--fcc-editable-region--
```
