---
id: 6853f2da9b8ed891e9fdc5fc
title: Passo 23
challengeType: 20
dashedName: step-23
---

# --description--

Agora vamos exibir as informações do remetente e do destinatário. Para mostrar quem enviou e recebeu o email, adicione duas instruções de impressão ao método `display_full_email` neste formato:

- `From: sender` onde `sender` é substituído pelo nome do remetente 
- `To: receiver` onde `receiver` é substituído pelo nome do destinatário

# --hints--

Você deve usar f-strings para imprimir o nome do remetente usando `From: {self.sender.name}`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'From: {self.sender.name}')")`));
    }
})
```

Você deve usar f-strings para imprimir o nome do destinatário usando `To: {self.receiver.name}`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'To: {self.receiver.name}')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
