---
id: 6853f4584f508e95cc7a9f4d
title: Passo 30
challengeType: 20
dashedName: step-30
---

# --description--

Agora você vai criar um método para listar todos os emails na caixa de entrada. Este método deve lidar com o caso em que a caixa de entrada está vazia e exibir uma lista numerada de emails quando houver emails presentes.

Adicione um método chamado `list_emails` à sua classe `Inbox` que recebe `self` como parâmetro. Primeiro, crie uma declaração `if` para verificar se a caixa de entrada está vazia testando a lista `emails`. Se estiver vazia, imprima a mensagem `Your inbox is empty.\n`. Adicione uma declaração `return` para sair do método.

# --hints--

Você deve definir um método chamado `list_emails` na classe `Inbox` que recebe `self` como parâmetro.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").has_args("self")`))
    }
})
```

Você deve verificar se `self.emails` está vazio usando `if not self.emails`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }
})
```

Você deve imprimir a mensagem `Your inbox is empty.\n` se não houver e-mails.

```js
({
    test: () => 
    { 
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[0].is_equivalent("print('Your inbox is empty.\\\\n')")`))
    }
})
```

Você deve ter uma declaração `return` no final da declaração `if`.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[1].is_equivalent("return")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    
--fcc-editable-region--

```
