---
id: 6853f90fd03acc9fd75ff860
title: Passo 46
challengeType: 20
dashedName: step-46
---

# --description--

Os usuários devem ser capazes de verificar sua caixa de entrada, ler e-mails e excluir e-mails diretamente através do objeto User deles.

Para verificar a caixa de entrada, adicione um método chamado `check_inbox` à classe `User`. Este método deve imprimir um cabeçalho personalizado com o nome do usuário e então exibir todos os seus e-mails.

O cabeçalho deve ser formatado como: `\nName's Inbox:`, onde `Name` é substituído pelo nome do usuário.

Após imprimir o cabeçalho, chame o método `list_emails()` na caixa de entrada do usuário.

# --hints--

Você deve definir um método chamado `check_inbox` na classe `User` que recebe `self` como parâmetro.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("check_inbox").has_args("self")`))
    }

})
```

Dentro de `check_inbox`, você deve imprimir um cabeçalho que inclua o nome do usuário no formato `\n[user]'s Inbox:`.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("check_inbox").find_body()[0].is_equivalent('print(f"\\\\n{self.name}\\\'s Inbox:")')`))
    }
})
```

Dentro de `check_inbox`, você deve chamar `self.inbox.list_emails()` para exibir os emails.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("check_inbox").find_body()[1].is_equivalent("self.inbox.list_emails()")`))
    }
})

```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)
        print(f'Email sent from {self.name} to {receiver.name}!\n')

--fcc-editable-region--
    
--fcc-editable-region--

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        del self.emails[actual_index]
        print('Email deleted.\n')

```
