---
id: 6853f34ff93594934f6636d9
title: Passo 27
challengeType: 20
dashedName: step-27
---

# --description--

Vamos adicionar uma representação em string à nossa classe `Email` para que possamos exibir resumos breves de e-mails.

Adicione um método `__str__` à classe `Email` que receba `self` como parâmetro.

# --hints--

Você deve definir um método `__str__` na classe `Email`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__str__")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

--fcc-editable-region--
    
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
