---
id: 6853f4ded9f84896ff182dd8
title: Passo 31
challengeType: 20
dashedName: step-31
---

# --description--

ApĂłs a verificaĂ§ĂŁo da caixa de entrada vazia, imprima a mensagem `\nYour Emails:`. Depois disso, itere sobre todos os emails na caixa de entrada usando um loop `for` com enumeraĂ§ĂŁo. Esta Ă© a sintaxe para enumeraĂ§ĂŁo com um nĂșmero inicial: `enumerate(iterable, start=start_number)`.

Use enumeraĂ§ĂŁo para numerar os emails comeĂ§ando a partir de `1`. Use `i` (o Ă­ndice) e `email` (o objeto email) como as variĂĄveis de iteraĂ§ĂŁo.

Dentro do loop, imprima uma f-string com o Ă­ndice da iteraĂ§ĂŁo seguido por um `.`, depois um espaĂ§o e a representaĂ§ĂŁo em string do objeto email neste formato:

```md
i. string_representation
```

Onde `i` Ă© o Ă­ndice e `string_representation` Ă© a representaĂ§ĂŁo do objeto email.

# --hints--

VocĂȘ deve imprimir o cabeĂ§alho `\nYour Emails:` fora da declaraĂ§ĂŁo `if`.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_body()[1].is_equivalent("print('\\\\nYour Emails:')")`))
    }
})
```

VocĂȘ deve ter um loop `for` que itera sobre `enumerate(self.emails, start=1)`.

```js
({
    test: () => 
    {
        runPython(`
_iter = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_iter()
assert _iter.is_equivalent("enumerate(self.emails, start=1)") or _iter.is_equivalent("enumerate(self.emails, 1)")
`)
    }
})
```

Seu loop `for` deve usar `i` e `email` para iterar sobre `enumerate(self.emails, start=1)`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_vars().is_equivalent("i, email")`))
    }
})
```

Dentro do loop, vocĂȘ deve imprimir cada resumo de email com um prefixo numerado seguido da string formatada do mĂ©todo `__str__`. NĂŁo esqueĂ§a de adicionar um `.` apĂłs o nĂșmero.

```js
({
    test: () => 
    {
        assert(runPython(`
        cond = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_body()
        cond.is_equivalent("print(f'{i}. {str(email)}')") or cond.is_equivalent("print(f'{i}. {email}')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
        
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
--fcc-editable-region--
        
--fcc-editable-region--
```
