---
id: 6853f4ded9f84896ff182dd8
title: Passo 31
challengeType: 20
dashedName: step-31
---

# --description--

Após a verificação da caixa de entrada vazia, imprima a mensagem `\nYour Emails:`. Depois disso, itere sobre todos os emails na caixa de entrada usando um loop `for` com enumeração. Esta é a sintaxe para enumeração com um número inicial: `enumerate(iterable, start=start_number)`.

Use enumeração para numerar os emails começando a partir de `1`. Use `i` (o índice) e `email` (o objeto email) como as variáveis de iteração.

Dentro do loop, imprima uma f-string com o índice da iteração seguido por um `.`, depois um espaço e a representação em string do objeto email neste formato:

```md
i. string_representation
```

Onde `i` é o índice e `string_representation` é a representação do objeto email.

# --hints--

Você deve imprimir o cabeçalho `\nYour Emails:` fora da declaração `if`.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_body()[1].is_equivalent("print('\\\\nYour Emails:')")`))
    }
})
```

Você deve ter um loop `for` que itera sobre `enumerate(self.emails, start=1)`.

```js
({
    test: () => 
    {
        runPython(`
_iter = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_iter()
assert _iter.is_equivalent("enumerate(self.emails, start=1)") or _iter.is_equivalent("enumerate(self.emails, 1)")
`)
    }
})
```

Seu loop `for` deve usar `i` e `email` para iterar sobre `enumerate(self.emails, start=1)`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_vars().is_equivalent("i, email")`))
    }
})
```

Dentro do loop, você deve imprimir cada resumo de email com um prefixo numerado seguido da string formatada do método `__str__`. Não esqueça de adicionar um `.` após o número.

```js
({
    test: () => 
    {
        assert(runPython(`
        cond = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_body()
        cond.is_equivalent("print(f'{i}. {str(email)}')") or cond.is_equivalent("print(f'{i}. {email}')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        
--fcc-editable-region--
```
