---
id: 68594df7779a164e6a61e8ae
title: Passo 32
challengeType: 20
dashedName: step-32
---

# --description--

A caixa de entrada precisa de um método para ler um email específico. Quando um usuário quiser ler um email, ele especificará qual número do email (começando pelo índice 0) deseja ver, e o método exibirá o conteúdo completo do email.

Adicione um método chamado `read_email` à sua classe `Inbox` que recebe um parâmetro `index`. Primeiro, verifique se a caixa de entrada está vazia e imprima a mensagem `Inbox is empty.\n` se estiver. Adicione uma instrução `return` depois disso para sair do método.

# --hints--

Você deve definir um método chamado `read_email` na classe `Inbox` que recebe `self` e `index` como parâmetros.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").has_args("self, index")`));
    }
})
```

Você deve verificar se `self.emails` está vazio usando `if not self.emails`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`));
    }
})
```

Você deve imprimir `Inbox is empty.\n` se não houver e-mails e retornar.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[0].is_equivalent("print('Inbox is empty.\\\\n')")`));
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[1].is_equivalent("return")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

--fcc-editable-region--
    
--fcc-editable-region--
```
