---
id: 6859400f8f7f872e50e1238c
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

A classe `User` precisa de um método para enviar e-mails para outros usuários. Ao enviar um e-mail, você criará um novo objeto `Email` e o adicionará à caixa de entrada do destinatário.

Crie um método chamado `send_email` na sua classe `User`. Este método deve receber parâmetros para `receiver`, `subject` e `body`. 

Dentro do método, crie um novo objeto `Email` com os seguintes valores de parâmetro e atribua-o a uma variável chamada `email`:

- `sender`: `self` (o usuário atual)
- `receiver`: `receiver` (o usuário que recebe o email)
- `subject`: `subject` (o assunto do email)
- `body`: `body` (o corpo do email)

# --hints--

Você deve definir um método `send_email` na classe `User`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").has_function("send_email")`))
    }

})
```

O método `send_email` deve ter os parâmetros: `self`, `receiver`, `subject` e `body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_args("self, receiver, subject, body")`))
    }

})
```

O método deve criar um objeto `Email` chamado `email` com os parâmetros corretos.

```js
({
  test: () => {
    assert(runPython(`
cond = _Node(_code).find_class("User").find_function("send_email").find_body()[0]
cond.is_equivalent("email = Email(sender=self, receiver=receiver, subject=subject, body=body)") or cond.is_equivalent("email = Email(self, receiver, subject, body)")
    `));
  }
});
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
```
