---
id: 68593ea22eac7e228ed585ce
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Os emails devem rastrear se foram lidos ou não. Adicione um atributo `read` ao método `__init__` e defina-o como `False` por padrão, já que novos emails começam como não lidos.

# --hints--

Você deve atribuir `False` a `self.read` no método `__init__`.

```js
({ 
    test: () => runPython(`
    assert _Node(_code).find_class("Email").find_function("__init__").find_variable("self.read").is_equivalent("self.read = False")
    `) 
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
```
