---
id: 685c1e32254e5993a5c8823e
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

Agora que você tem uma configuração inicial da classe `Email`, é hora de criar usuários que podem enviar e receber emails. Cada usuário terá um nome e poderá realizar operações de email.

Crie uma classe `User` com um método `__init__` que recebe `self` e `name` como parâmetros. Dentro da classe, atribua o parâmetro `name` a `self.name`.

# --hints--

Você deve criar uma classe `User`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_class("User")`))
    }

})
```

A classe `User` deve ter um método `__init__` que recebe `self` e `name` como parâmetros.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").has_args("self, name")`))
    }

})
```

Você deve atribuir o parâmetro `name` a `self.name`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.name").is_equivalent("self.name = name")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--

--fcc-editable-region--
```
