---
id: 685cdfd1932d8ac6ad986813
title: Passo 28
challengeType: 20
dashedName: step-28
---

# --description--

Dentro do método `__str__`, você mostrará se um email foi lido ou não. 

Expressões condicionais permitem que você atribua um valor com base em uma condição em uma única linha.

Aqui está um exemplo:

```py
x = 10
y = 'Even' if x % 2 == 0 else 'Odd' # y will be even
```

Dentro do método, antes, use expressão condicional para atribuir a string `Read` a uma variável `status` se o email estiver lido e `Unread` se não estiver. 

# --hints--

O método `__str__` deve criar uma variável `status` que use uma expressão condicional.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__str__").has_stmt("status = 'Read' if self.read else 'Unread'")`));
    }
})
```

O método `__str__` deve retornar as informações do email com o status incluído.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__str__").has_variable("status")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

--fcc-editable-region--
    def __str__(self):
        pass
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
