---
id: 685c1eba966c6396850736ea
title: Passo 13
challengeType: 20
dashedName: step-13
---

# --description--

Embora os usuários possam enviar e-mails e ter caixas de entrada, precisamos de uma classe dedicada para gerenciar as operações da caixa de entrada de forma eficiente. 

A classe `Inbox` armazenará uma lista de emails e fornecerá métodos para adicionar novos emails, listar todos os emails e gerenciá-los.

Crie uma nova classe chamada `Inbox` com um método `__init__` que inicializa uma lista vazia chamada `emails`.

# --hints--

Você deve criar uma classe `Inbox`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_class("Inbox")`));
    }
})
```

A classe `Inbox` deve ter um método `__init__` que receba `self` como parâmetro.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").has_args("self")`));
    }
})
```

Você deve inicializar `self.emails = []` no método `__init__` da classe `Inbox`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").find_variable("self.emails").is_equivalent("self.emails = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--

--fcc-editable-region--
```
