---
id: 688754f3494b17e0b352d67d
title: Passo 17
challengeType: 20
dashedName: step-17
---

# --description--

Agora é hora de testar nosso sistema completo de email! Vamos criar alguns usuários e ver a funcionalidade de email em ação.

Crie dois objetos User: `alice` com nome `"Alice"` e `bob` com nome `"Bob"`. Isso demonstrará como os usuários podem ser criados e interagir entre si.

# --hints--

Você deve criar um objeto User chamado `alice` com o nome `"Alice"`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("alice").is_equivalent('alice = User("Alice")')`));
    }
})
```

Você deve criar um objeto User chamado `bob` com o nome `"Bob"`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("bob").is_equivalent('bob = User("Bob")')`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--

--fcc-editable-region--
```
