---
id: 68875554c6c32de788afc64e
title: Passo 21
challengeType: 20
dashedName: step-21
---

# --description--

Os usuários precisam de uma maneira de ler seus emails corretamente.

Adicione um método `display_full_email` à classe `Email` que recebe apenas `self` como parâmetro. Dentro desse método, chame o método `mark_as_read()` para marcar o email como lido quando alguém o visualizar.

Este método seria usado para exibir o conteúdo completo do email de uma forma bem formatada.

# --hints--

Você deve definir um método `display_full_email` na classe `Email`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").has_function("display_full_email")`));
    }
})
```

O método `display_full_email` deve receber apenas `self` como parâmetro.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_args("self")`));
    }
})
```

O método deve chamar `self.mark_as_read()` para marcar o email como lido.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("self.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True
--fcc-editable-region--
    
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
