---
id: 688754d6d59578df1313cb3c
title: Passo 14
challengeType: 20
dashedName: step-14
---

# --description--

Agora é hora de conectar as classes User e Inbox para que os emails possam realmente ser entregues! 

Os usuários precisam ter objetos Inbox adequados em vez de listas simples e o método `send_email` deve entregar os emails na caixa de entrada do destinatário.

Atualize o `User` class `__init__` de forma que cada usuário receba um objeto Inbox real em vez de uma lista vazia criada anteriormente. 

# --hints--

Você deve atualizar a classe User para inicializar `self.inbox = Inbox()`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = Inbox()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []
--fcc-editable-region--

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

class Inbox:
    def __init__(self):
        self.emails = []

```
