---
id: 68b9a011daf902ff1eba40a8
title: Passo 20
challengeType: 20
dashedName: step-20
---

# --description--

Agora que o email foi recebido, remova as criações de objetos, chamadas de métodos e declarações de impressão.

# --hints--

Você não deve ter `alice = User("Alice")` no seu código.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_stmt("alice = User(\\"Alice\\")")`));
    }
})
```

Você não deve ter `bob = User("Bob")` no seu código.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_stmt("bob = User(\\"Bob\\")")`));
    }
})
```

Você não deve ter `alice.send_email(bob, "Hello", "Hi Bob, how are you?")` no seu código.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("alice.send_email(bob, \\"Hello\\", \\"Hi Bob, how are you?\\")")`));
    }
})
```

Você não deve ter `print(len(bob.inbox.emails))` no seu código.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(len(bob.inbox.emails))")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
alice = User("Alice")
bob = User("Bob")

alice.send_email(bob, "Hello", "Hi Bob, how are you?")

print(len(bob.inbox.emails))
--fcc-editable-region--
```
