---
id: 68b930ee0e7b8fb81895e59a
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Agora você vai adicionar um método para marcar um email como lido. Quando alguém abrir um email, você quer mudar seu status de não lido para lido. Esse método será simples - ele só precisa definir o atributo `read` como `True`.

Adicione um método chamado `mark_as_read` à sua classe `Email`. Este método deve receber apenas `self` como parâmetro e definir o atributo `read` como `True`.

# --hints--

Você deve definir um método chamado `mark_as_read` na classe `Email` que deve receber apenas `self` como parâmetro.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").has_args("self")`))
    }

})
```

O método deve definir `self.read` como `True`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").find_variable("self.read").is_equivalent("self.read = True")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
```
