---
id: 68b9366fae8b13c1278761b1
title: Passo 15
challengeType: 20
dashedName: step-15
---

# --description--

Sua caixa de entrada precisa de uma forma de receber novos emails. Quando alguém envia um email para um usuário, ele deve ser adicionado à sua caixa de entrada.

Adicione um método chamado `receive_email` à sua classe `Inbox` que recebe `self` e um objeto de email `email` como parâmetro. Dentro do corpo do método, adicione o `email` à lista `emails` usando o método `append`.

# --hints--

Você deve definir um método chamado `receive_email` na classe `Inbox` que recebe `self` e `email` como parâmetros.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_args("self, email")`))
    }
})
```

Você deve adicionar o `email` a `self.emails`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_stmt("self.emails.append(email)")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--
class Inbox:
    def __init__(self):
        self.emails = []

    
--fcc-editable-region--
```
