---
id: 68b9350183ff05bd106c79e4
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

Agora, teste o método `mark_as_read` que você criou na classe `Email`. Use o método no objeto `Email` existente `email_obj` para alterar o status. Após marcar o email como lido, imprima o atributo `read` de `email_obj` para confirmar a alteração.

# --hints--

Você deve chamar o método `mark_as_read` no `email_obj` para marcar o email como lido.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`))
    }

})
```

Você deve imprimir o atributo `read` do `email_obj` após chamar o método `mark_as_read`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).is_ordered("email_obj.mark_as_read()", "print(email_obj.read)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)

--fcc-editable-region--
```
