---
id: 68b936bf4f47d6c29a669f99
title: Passo 16
challengeType: 20
dashedName: step-16
---

# --description--

Ainda dentro do método `send_email`, chame o método `receive_email` da `inbox` do `receiver` para adicionar o email à caixa de entrada do receptor.

# --hints--

Você deve adicionar `receiver.inbox.receive_email(email)` ao método `send_email`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_stmt("receiver.inbox.receive_email(email)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()
        
--fcc-editable-region--
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)
```
