---
id: 68bf0f09b7ec33be465c0c35
title: Passo 29
challengeType: 20
dashedName: step-29
---

# --description--

No final do método `__str__`, retorne o resumo do email neste formato:

```py
[status] From: sender | Subject: subject.
```

Onde `status` é o status do email, `sender` é o nome do remetente e `subject` é o assunto do email.

# --hints--

O método `__str__` deve retornar `f"[{status}] From: {self.sender.name} | Subject: {self.subject]"`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__str__").has_return("f\\"[{status}] From: {self.sender.name} | Subject: {self.subject}\\"")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

--fcc-editable-region--
    def __str__(self):
        status = 'Read' if self.read else 'Unread'
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
