---
id: 68c18b1d0fce081f85312d97
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

У попередньому уроці ви дізналися, що атрибут із префіксом одного підкреслення призначений для внутрішнього використання за конвенцією.

Змініть обидва атрибути `name` і `level` на `_name` і `_level`, оскільки їх не слід змінювати поза межами їхнього класу.

Зверніть увагу, що це не забороняє доступ або зміну атрибутів поза межами класу. Також у Python завжди є спосіб отримати доступ до приватних атрибутів (із подвійним підкресленням на початку).

# --hints--

Ваш метод `__init__` має встановлювати `self._name` у значення `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._name = name")`)) })
```

Ваш метод `__init__` має встановлювати `self._level` у значення `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level

charlie_brown = Employee('Charlie Brown', 'trainee')
--fcc-editable-region--
```
