---
id: 68c15a2a1f0e70ca7154aba5
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому практичному занятті ви створите систему відстеження зарплат для працівників.

Почніть зі створення класу з іменем `Employee`. Усередині нього створіть метод `__init__` з параметрами `self`, `name` та `level`. У методі `__init__` призначте `name` і `level` атрибутам екземпляра з такими ж іменами.

# --hints--

У вас має бути клас з іменем `Employee`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Employee")`)) })
```

Ваш клас `Employee` повинен мати метод `__init__` з трьома параметрами `self`, `name` та `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_args("self, name, level")`)) })
```

Ваш метод `__init__` має призначати `self.name` значенню `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

Ваш метод `__init__` має призначати `self.level` значенню `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
