---
id: 68c16924498728d259e4781c
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Тепер створіть екземпляр класу `Employee`, передаючи рядки `Charlie Brown` і `trainee`. Призначте цей екземпляр змінній з ім’ям `charlie_brown`.

# --hints--

У вас має бути змінна з ім’ям `charlie_brown`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("charlie_brown")`)) })
```

Ви маєте призначити `Employee('Charlie Brown', 'trainee')` змінній `charlie_brown`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("charlie_brown").is_equivalent("charlie_brown = Employee('Charlie Brown', 'trainee')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level


--fcc-editable-region--
```
