---
id: 68c27fa1697ceb7c578b2c9f
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Декоратор `@property` використовується в Python, щоб перетворити метод на властивість. Зазвичай його застосовують для визначення геттерів — методів, які використовують для отримання значення атрибута:

```py
class Person:
    def __init__(self, name):
        self._name = name

    @property
    def name(self):
        return self._name

p = Person('Alice')
print(p.name)  # Alice
```

Створіть метод з ім’ям `name` з параметром `self` і задекоруйте його `@property`. Усередині методу поверніть `self._name`.
    
# --hints--

Ваш клас `Employee` має містити метод `name` з параметром `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_args("self")`)) })
```

Ваш метод `name` має бути задекорований `@property`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_decorators("property")`)) })
```

Ваш метод `name` має повертати `self._name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_return("self._name")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
