---
id: 68c27eb292bfec6b9402d58f
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Додайте метод `__str__` до класу `Employee`. Нехай він повертає f-рядок у форматі `name: level`, замінюючи `name` і `level` на відповідні атрибути.

Після цього виведіть `charlie_brown` у консоль.

# --hints--

Ваш клас `Employee` має містити метод `__str__` з параметром `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_args("self")`)) })
```

Ваш метод `__str__` має повертати f-рядок у форматі `name: level`, замінюючи `name` і `level` на значення `self._name` і `self._level` відповідно.

```js
({ test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    empl = Employee("Frank", "dreamer")
    print(empl)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "Frank: dreamer" in output
`) })
```

Ви маєте вивести `charlie_brown` у консоль.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    

charlie_brown = Employee('Charlie Brown', 'trainee')

--fcc-editable-region--
```
