---
id: 68c281cecda94ba2c6838e17
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Тепер, коли ви визначили гетер для `name`, ви можете отримати доступ до атрибута `_name` через властивість `name`. Тож виведіть `charlie_brown.name` у консоль.

# --hints--

Ви повинні вивести `charlie_brown.name` у консоль.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown.name)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    @property
    def name(self):
        return self._name

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
