---
id: 68c2d31a86d5672ca49d4521
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Тепер час додати валідацію до методу `__init__`. На початку методу створіть інструкцію `if`, яка перевіряє, чи не є `name` або `level` екземплярами `str`.

Всередині інструкції `if` викличте `TypeError` з повідомленням `'name' and 'level' attribute must be of type 'str'.`

# --hints--

У вашому методі `if` має бути інструкція `__init__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0]`)) })
```

Ваша інструкція `if` має перевіряти, чи не є `name` або `level` екземплярами `str`.

```js
({ test: () => runPython(`
  cond = _Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0].find_conditions()[0]
  conditions = [
    "not (isinstance(name, str) and isinstance(level, str))",
    "not isinstance(name, str) or not isinstance(level, str)",
    "not (isinstance(level, str) and isinstance(name, str))",
    "not isinstance(level, str) or not isinstance(name, str)",
  ]
  assert any(cond.is_equivalent(c) for c in conditions)
`) })
```

Ваша інструкція `if` має викликати `TypeError` з повідомленням `'name' and 'level' attribute must be of type 'str'.`

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError("\\'name\\' and \\'level\\' attribute must be of type \\'str\\'.")')`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
--fcc-editable-region--
    def __init__(self, name, level):
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
