---
id: 68c28d659d345270d4a61084
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Метод `__repr__` — це спеціальний метод, який має повертати рядкове представлення об’єкта, що може бути використане для його створення.

Наприклад, метод `__repr__` класу `Employee('Charlie', 'developer')` має повертати рядок `Employee('Charlie', 'developer')`, який є тим самим рядком, що використовується для створення об’єкта.

Додайте до свого класу `Employee` метод `__repr__` з параметром `self` і зробіть так, щоб він повертав рядок, який можна використати для створення об’єкта.

# --hints--

Ваш клас `Employee` має мати метод `__repr__` з параметром `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__repr__").has_args("self")`)) })
```

Ваш метод `__repr__` має повертати рядок, який використовується для створення об’єкта.

```js
({ test: () => runPython(`
  empl = Employee("Frank", "dreamer")
  assert repr(empl) == 'Employee("Frank", "dreamer")' or repr(empl) == "Employee('Frank', 'dreamer')"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
