---
id: 68c7f7012a700243eff1cbc0
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

Рівень, який використовується для створення працівника, має бути обраний серед конкретних рівнів. Ви використаєте `_base_salaries` для перевірки рівня та встановлення правильної зарплати працівнику.

Після вашої існуючої `if` інструкції створіть ще одну `if`, яка перевіряє, чи `level` не входить до `Employee._base_salaries`.

Всередині нової `if` інструкції викличте `ValueError` з повідомленням `Invalid value '{level}' for 'level' attribute.`, де `{level}` має бути замінений на значення аргументу `level`.

# --hints--

Коли `level` не входить до `Employee._base_salaries`, ви маєте викликати `ValueError` з повідомленням `Invalid value '{level}' for 'level' attribute.`, де `{level}` має бути замінений на значення аргументу `level`.

```js
({ test: () => runPython(`
  try:
    Employee("Frank", "dreamer")
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
    assert False, "Expected to raise ValueError with invalid level"
`) })
```

Ви не повинні викликати жодного виключення, коли `level` входить до `Employee._base_salaries`.

```js
({ test: () => runPython(`
  levels = [
    "trainee",
    "junior",
    "mid-level",
    "senior"
  ]
  for level in levels:
    try:
      Employee("Frank", level)
    except Exception:
      assert False, f"Expected not to raise ValueError with level '{level}'"
    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }
--fcc-editable-region--
    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
