---
id: 68c9cab4b1118da59eecfc56
title: Крок 31
challengeType: 20
dashedName: step-31
---

# --description--

Після існуючих інструкцій `if` створіть ще одну, щоб викликати `ValueError`, коли `new_level` вже є вибраним рівнем.

Зверніть увагу, що `_level` ще не існує під час ініціалізації, тому використовуйте `hasattr(self, '_level')` для перевірки його існування перед порівнянням. Це уникне `AttributeError` при першому створенні об’єкта.

Для повідомлення використовуйте `'{level}' is already the selected level.`, де `{level}` має бути замінений на поточний рівень.

# --hints--

У вас має бути третя інструкція `if` всередині вашого сеттера `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2]`)) })
```

Коли `new_level` дорівнює `self.level`, ви маєте викликати `ValueError` з повідомленням `'{level}' is already the selected level.`, де `{level}` має бути замінений на поточний рівень.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.level = "trainee"
  except ValueError as e:
    assert str(e) == "'trainee' is already the selected level."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

Ваша інструкція `if` має використовувати `hasattr(self, '_level')` для перевірки існування `_level` перед порівнянням.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2].find_conditions()[0].is_equivalent("hasattr(self, '_level') and new_level == self.level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level
    

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
