---
id: 68ca758f8160b11757f877ae
title: Крок 32
challengeType: 20
dashedName: step-32
---

# --description--

Нарешті, створіть четверту інструкцію `if`, яка піднімає `ValueError` з повідомленням `Cannot change to lower level.`, коли базова зарплата нового рівня менша за базову зарплату поточного рівня.

Використовуйте `hasattr(self, '_level')`, щоб уникнути `AttributeError` під час ініціалізації.

# --hints--

У вас має бути четверта інструкція `if` всередині вашого сеттера `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[3]`)) })
```

Ваша інструкція `if` має використовувати `hasattr(self, '_level')`, щоб перевірити, чи існує `_level`, перш ніж порівнювати.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[3].find_conditions()[0].is_equivalent("hasattr(self, '_level') and Employee._base_salaries[new_level] < Employee._base_salaries[self.level]")`)) })
```

Коли `new_level` менша за `self.level`, ви маєте підняти `ValueError` з повідомленням `Cannot change to lower level.`.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'junior')
  try:
    emp.level = "trainee"
  except ValueError as e:
    assert str(e) == "Cannot change to lower level."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

Ви не повинні піднімати жодних виключень, коли `new_level` більша за `self.level`.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'junior')
  try:
    emp.level = "senior"
  except Exception:
    assert False, "Expected not to raise exception with valid new_level"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
        if hasattr(self, '_level') and new_level == self.level:
            raise ValueError(f"'{self.level}' is already the selected level.")
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level
    
    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
