---
id: 6a3e6c0eb5d050ce07781709
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

如果 `if` 條件評估為 `False`（表示你確實有足夠的蛋），函式會跳過 `if` 區塊並繼續執行剩下的程式碼。

現在你需要處理你的資料並傳回一個新的計算值。

在你的 `if` 區塊外（與 `if` 關鍵字相同的縮排層次），使用 `print()` 呼叫和 f-string 列印 `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'`。

接著，在下一行，透過從 `available_eggs` 減去 `eggs_to_use` 來 `return` 剩餘的蛋數。

# --hints--

你的 `use_eggs` 函式中應該有一個列印 `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'` 的列印述語。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_stmt("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')")
`)
})
```

你應該傳回透過從 `available_eggs` 減去 `eggs_to_use` 計算出的剩餘蛋數。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_return("available_eggs - eggs_to_use")
`)
})
```

你應該在 `print` 述語之後傳回 `available_eggs - eggs_to_use`。

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs")
assert node.is_ordered("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')", "return available_eggs - eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

--fcc-editable-region--
    
--fcc-editable-region--
```
