---
id: 6a3e71a0a3a44260e675b157
title: 第 12 步驟
challengeType: 20
dashedName: step-12
---

# --description--

為了解開為什麼你的蛋數量沒有改變並且實際上煮早餐的謎團，創建一個新的函式（程式）。把這個函式想成你的食譜，它會呼叫其他函式來完成工作。

在你的 `print(available_eggs)` 述語下面，定義一個空的函式叫做 `make_fried_egg`，它只接受 `available_eggs` 作為參數。

目前在函式主體內使用 `pass` 關鍵字。

# --hints--

你應該定義一個名為 `make_fried_egg` 的函式。

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("make_fried_egg")
`)
})
```

你的 `make_fried_egg` 函式應該有 `available_eggs` 作為參數。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").has_args("available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

--fcc-editable-region--

--fcc-editable-region--
```
