---
id: 6a3e6fd08125dfa615cb431b
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

現在你的 `use_eggs` 函式已經完成，可以開始測試。當呼叫需要**參數**的函式時，你會將具體的值放在小括號內。這些值稱為**引數**。

以下是如何使用引數呼叫函式的範例：

```py
def calculate_sum(a, b):
    print(a + b)
calculate_sum(3, 1) # Output: 4
```

在你的程式碼底部，呼叫你的 `use_eggs` 函式，並傳入兩個引數：全域的 `available_eggs` 變數和數字 `1`（代表你想使用的蛋的數量）。

# --hints--

你應該呼叫 `use_eggs` 函式，並傳入 `available_eggs` 和 `1` 作為引數。

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("use_eggs(available_eggs, 1)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

--fcc-editable-region--

--fcc-editable-region--
```
