---
id: 6a3e70b94f67d9ee3ecd8486
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

你剛剛呼叫了 `use_eggs(available_eggs, 1)` 來使用一顆蛋。那代表你廚房裡應該只剩下 `0` 顆蛋，對吧？

你很快就會知道答案！在你的程式碼底部，使用 `print()` 函式列印 `available_eggs` 變數的值。

當你執行程式碼時，你會發現一件令人驚訝的事：輸出仍然是 `1`！為什麼蛋的數量沒有改變？你會在接下來的幾個步驟中探索這個謎團。

# --hints--

你應該在程式碼的最後列印 `available_eggs` 變數。

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("print(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
--fcc-editable-region--

--fcc-editable-region--
```
