---
id: 6a3e73eabddd614f91cc235a
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

現在你要來解開之前的謎團！當你的 `use_eggs` 函式傳回剩餘的蛋數時，你並沒有把那個新值存到任何地方。它就這樣消失在虛空中了！

要真正更新你的庫存，你需要透過指定一個變數來接收那個傳回的值。

在你的 `make_fried_egg` 函式裡，使用 `if/else` 述語來檢查 `has_enough_eggs` 變數。

- 如果 `has_enough_eggs` 是 `True`，將你的 `available_eggs` 變數重新指定為 `use_eggs(available_eggs, 1)` 的結果，並列印訊息 `Made a fried egg. Yummy!`。

- 否則，在 `else` 區塊中列印訊息 `Could not make a fried egg. Not enough eggs!`。

# --hints--

你應該有一個 `if` 述語，並使用 `has_enough_eggs` 作為條件。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_conditions()[0].is_equivalent("has_enough_eggs")
`)
})
```

在你的 `if` 區塊中，你應該將 `available_eggs` 重新指定為 `use_eggs(available_eggs, 1)` 的結果。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].find_variables("available_eggs")[0].is_equivalent("available_eggs = use_eggs(available_eggs, 1)")
`)
})
```

當 `has_enough_eggs` 是 `True` 時，你應該列印訊息 `Made a fried egg. Yummy!`。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].has_call("print('Made a fried egg. Yummy!')")
`)
})
```

你應該有一個 `else` 區塊，當 `has_enough_eggs` 是 `False` 時列印訊息 `Could not make a fried egg. Not enough eggs!`。

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg").find_if("has_enough_eggs")
conditions = node.find_conditions()
assert len(conditions) == 2 and conditions[1].is_empty()
assert node.find_bodies()[1].has_call("print('Could not make a fried egg. Not enough eggs!')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

--fcc-editable-region--
    
--fcc-editable-region--
```
