---
id: 6a3e7284e89ce3367e709331
title: 第 13 步驟
challengeType: 20
dashedName: step-13
---

# --description--

在煎蛋之前，你需要先檢查是否有足夠的食材。你可以透過比較可用的蛋數量和食譜所需的蛋數量來完成這個檢查。

在你的 `make_fried_egg` 函式（程式）中，將 `pass` 關鍵字替換成一個名為 `has_enough_eggs` 的變數。將它指定為檢查 `available_eggs` 是否大於或相等於 `1` 的結果。

# --hints--

你應該將 `has_enough_eggs` 指定為檢查 `available_eggs` 是否大於或相等於 `1` 的結果。

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg").find_variables("has_enough_eggs")[0]
assert node.is_equivalent("has_enough_eggs = available_eggs >= 1") or node.is_equivalent("has_enough_eggs = 1 <= available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
