---
id: 6a3eaa2caac22f67811fe40f
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

真相時刻到了！你已成功用 `make_fried_egg` 函式（程式）傳回的新值更新了全域的 `available_eggs` 變數。

最後再檢查一次你的廚房庫存，確認雞蛋數量確實減少了。

在你的程式碼最後，再次呼叫 `check_kitchen_stock()` 函式（程式）。當你執行程式碼時，你應該會看到廚房現在有 `0` 顆雞蛋！

恭喜你成功完成廚房庫存追蹤工作坊！

# --hints--

你應該在程式碼最後第二次呼叫 `check_kitchen_stock()` 函式（程式）。

```js
({
  test: () => runPython(`
node = _Node(_code)
assert len(node.find_calls("check_kitchen_stock")) == 2
assert node.is_ordered("available_eggs = make_fried_egg(available_eggs)", "check_kitchen_stock()")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

available_eggs = make_fried_egg(available_eggs)

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

available_eggs = make_fried_egg(available_eggs)

check_kitchen_stock()
```
