---
id: 6a3e58891e1ff13d07f4ab48
title: 第 5 步
challengeType: 20
dashedName: step-5
---

# --description--

Python 會讀取 `def` 區塊並記住該函式（程式），但不會立即執行其中的程式碼。

要執行程式碼並看到輸出，你需要**呼叫**（或喚起）該函式（程式）。方法是輸入函式（程式）名稱，後面接著小括號。

例如：

```py
my_function()
```

在你的程式碼底部、函式（程式）外，呼叫你的 `check_kitchen_stock` 函式（程式）。

# --hints--

你應該在程式碼結尾呼叫 `check_kitchen_stock` 函式（程式）。確保它沒有縮排，這樣它才能在全域範圍執行。

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("check_kitchen_stock()")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

--fcc-editable-region--

--fcc-editable-region--
```
