---
id: 6a3e524fc3394ef52ed5b797
title: 第 3 步
challengeType: 20
dashedName: step-3
---

# --description--

函式（程式）有自己的區域範圍。這表示在函式（程式）內創建的變數就像廚房儲藏室裡的食材：它們只能在該特定函式（程式）內被存取。

不過，函式（程式）並非完全孤立。它們可以讀取定義在其範圍外的全域變數。它們可以使用這些變數來進行計算、做出決策，或創建新的區域變數。

在你的 `check_kitchen_stock` 函式（程式）內，將 `pass` 關鍵字替換成一個名為 `total_items` 的新變數，並將 `available_eggs`、`available_flour` 和 `available_sugar` 的總和指定給它。

# --hints--

你應該在你的 `check_kitchen_stock` 函式（程式）內定義一個名為 `total_items` 的變數。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_variable("total_items")
`)
})
```

你的 `total_items` 變數應該是 `available_eggs`、`available_flour` 和 `available_sugar` 的總和。

```js
({
  test: () => runPython(`
variable = _Node(_code).find_function("check_kitchen_stock").find_variables("total_items")[0]
assert any(
  variable.is_equivalent(code)
  for code in [
    "total_items = available_eggs + available_flour + available_sugar",
    "total_items = available_flour + available_eggs + available_sugar",
    "total_items = available_sugar + available_flour + available_eggs",
    "total_items = available_eggs + available_sugar + available_flour",
    "total_items = available_flour + available_sugar + available_eggs",
    "total_items = available_sugar + available_eggs + available_flour"
  ]
)
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
--fcc-editable-region--
    pass
--fcc-editable-region--
```
