---
id: 6a3e5ccb99fd6053af9032a1
title: 第 6 步
challengeType: 20
dashedName: step-6
---

# --description--

既然你已經知道廚房裡有什麼，你就可以用那些項目來烹飪。為此，創建一個接受參數的新函式。把參數想像成你在呼叫函式時傳入的資料的佔位符。

這是一個接受兩個參數 `a` 和 `b`，並將它們的總和列印到終端機的函式。

```py
def calculate_sum(a, b):
    print(a + b)

calculate_sum(3, 1) # Output: 4
```

在你的 `check_kitchen_stock` 函式下面，定義一個名為 `use_eggs` 的空函式，該函式接受兩個參數：依序為 `available_eggs` 和 `eggs_to_use`。

請注意，如果你給參數取與全域變數（例如 `available_eggs`）完全相同的名稱，函式會使用它自己的區域參數 `available_eggs`，而不是全域變數。這稱為名稱遮蔽，因為區域變數會「遮蔽」全域變數。

目前在函式主體內使用 `pass` 關鍵字。

# --hints--

你應該定義一個名為 `use_eggs` 的函式。

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("use_eggs")
`)
})
```

你的 `use_eggs` 函式應該有 `available_eggs` 和 `eggs_to_use` 作為它的參數。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_args("available_eggs, eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

--fcc-editable-region--

--fcc-editable-region--
```
