---
id: 6a3e67fc5ef4045201c28219
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

在你的 `use_eggs` 函式（程式）中，你需要檢查是否有足夠的蛋來烹飪。你可以使用 `if` 述語來做到這點。

將 `pass` 關鍵字替換成一個 `if` 述語，檢查 `eggs_to_use` 是否大於（`>`）`available_eggs`。在 `if` 區塊內，列印字串 `The kitchen does not have enough eggs.`。

# --hints--

你應該有一個 `if` 述語，檢查 `eggs_to_use > available_eggs`。

```js
({
  test: () => runPython(`
  cond = _Node(_code).find_function("use_eggs").find_ifs()[0].find_conditions()[0]
  assert cond.is_equivalent("eggs_to_use > available_eggs") or cond.is_equivalent("available_eggs < eggs_to_use")`)
})
```

你應該在你的 `if` 述語內列印 `The kitchen does not have enough eggs.`。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").find_ifs()[0].has_call("print('The kitchen does not have enough eggs.')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
