---
id: 6a3e4ae26f2b7e4b5b026cf8
title: 第 2 步
challengeType: 20
dashedName: step-2
---

# --description--

現在是時候創建一個函式來檢查你廚房裡有什麼。在 Python 中，你使用 `def` 關鍵字來定義函式，接著是函式名稱、小括號 `()`，以及冒號 `:`。

以下是一個定義簡單函式的範例：

```py
def my_function():
    pass
```

定義一個名為 `check_kitchen_stock()` 的空函式。暫時在函式主體內使用 `pass` 關鍵字，這樣 Python 才不會發出語法錯誤。

記得 Python 是靠縮排來定義程式碼區塊的，所以務必縮排函式主體。

# --hints--

你應該定義一個名為 `check_kitchen_stock` 的函式。

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("check_kitchen_stock")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

--fcc-editable-region--

--fcc-editable-region--
```
