---
id: 6a3ea8a55317ba7a08e21d5b
title: 第 16 步
challengeType: 20
dashedName: step-16
---

# --description--

現在你終於可以回答之前的問題了：為什麼全域的蛋數會維持在 `1`？當函式傳回一個值時，它不會自動覆寫你現有的變數。你必須明確地儲存那個新值！

你在 `make_fried_egg` 函式內使用了 `available_eggs` 變數，但它並不是你程式碼頂端定義的那個變數。函式內的那個是區域變數，只在函式執行期間存在。當函式結束時，區域變數會被丟棄，而全域的 `available_eggs` 則保持不變。

要更新你實際的廚房庫存，請用 `make_fried_egg` 函式傳回的新蛋數覆寫全域的 `available_eggs` 變數。

在程式碼的最底部，呼叫 `make_fried_egg` 函式，並將 `available_eggs` 作為引數傳入，然後將結果指定回全域的 `available_eggs` 變數。

# --hints--

你應該將 `available_eggs` 變數重新指定為 `make_fried_egg(available_eggs)` 的結果。

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[-1].is_equivalent("available_eggs = make_fried_egg(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

--fcc-editable-region--

--fcc-editable-region--
```
