---
id: 6a3e2b9b7eeca084dd4fb857
title: 第 1 步
challengeType: 20
dashedName: step-1
---

# --description--

歡迎來到廚房！就像廚師依賴食譜來準備餐點一樣，你將依賴函式來建置軟體。在這個工作坊中，你將學習如何撰寫可重複使用的 Python 函式並掌握範圍，精確了解在任何時刻你的程式碼中有哪些材料（變數）可用。

要開始，你應該為你的廚房庫存定義一些基本資料。創建三個變數：`available_eggs`、`available_flour` 和 `available_sugar`，並分別指定它們的值為 `1`、`2` 和 `3`。

# --hints--

你應該定義一個名為 `available_eggs` 的變數，並指定它的值為 `1`。

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[0].is_equivalent("available_eggs = 1")
`)
})
```

你應該定義一個名為 `available_flour` 的變數，並指定它的值為 `2`。

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_flour")[0].is_equivalent("available_flour = 2")
`)
})
```

你應該定義一個名為 `available_sugar` 的變數，並指定它的值為 `3`。

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_sugar")[0].is_equivalent("available_sugar = 3")
`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
