---
id: 6a3e6ad6b9554e31dcc30cd6
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

如果你沒有足夠的雞蛋，就不應該嘗試烹飪！你需要一種方法來停止函式繼續執行。

你可以使用 `return` 關鍵字來做到這點。當函式遇到 `return` 述語時，它會立即停止執行，並將傳回的值返回到呼叫它的地方。

例如：

```py
def cook_dish():
    message = "The dish is ready to serve!"
    return message
    print("This line will never run.")

print(cook_dish())  # Output: The dish is ready to serve!
```

在你的 `if` 區段中，列印述語下面，使用 `return` 關鍵字傳回 `available_eggs` 的值。

# --hints--

你應該在你的 `if` 述語內傳回 `available_eggs`。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").find_ifs()[0].has_return("available_eggs")
`)
})
```

你應該在你的 `if` 區段中，`print` 述語之後傳回 `available_eggs`。

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs").find_ifs()[0]
assert node.is_ordered("print('The kitchen does not have enough eggs.')", "return available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
--fcc-editable-region--
        
--fcc-editable-region--
```
