---
id: 6a3e5475c35dbc138d77c6f4
title: 第 4 步
challengeType: 20
dashedName: step-4
---

# --description--

既然你已經計算出總數，接著應該顯示庫存。回想之前的課程，你可以使用 `print()` 函式（程式）搭配 **f-strings** 來在輸出中包含變數。

在你的 `check_kitchen_stock` 函式（程式）中，於 `total_items` 變數下方，加入以下四個完全相同的列印述語：

1. `print(f'The kitchen has {total_items} total items:')`
2. `print(f'- {available_eggs} eggs')`
3. `print(f'- {available_flour} flour')`
4. `print(f'- {available_sugar} sugar')`

# --hints--

你的函式（程式）中應該包含 `print(f'The kitchen has {total_items} total items:')` 這個述語。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'The kitchen has {total_items} total items:')")
`)
})
```

`print(f'The kitchen has {total_items} total items:')` 這個述語應該是你函式（程式）中第一個 `print()` 呼叫。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").find_calls("print")[0].is_equivalent("print(f'The kitchen has {total_items} total items:')")
`)
})
```

你的函式（程式）中應該包含 `print(f'- {available_eggs} eggs')` 這個述語。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_eggs} eggs')")
`)
})
```

你的函式（程式）中應該包含 `print(f'- {available_flour} flour')` 這個述語。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_flour} flour')")
`)
})
```

你的函式（程式）中應該包含 `print(f'- {available_sugar} sugar')` 這個述語。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_sugar} sugar')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
--fcc-editable-region--
    
--fcc-editable-region--
```
