---
id: 697fc395f3d85f14b05cd9ca
title: Wiederherstellen einer kohärenten Erzählung aus einem Array von Storyfragmenten
challengeType: 26
dashedName: lab-story-fragment-restoration
---

# --description--

In diesem Labor stellen Sie eine kohärente Erzählung aus einem beschädigten Array von Storyfragmenten wieder her.

Sie üben den Umgang mit Schleifen, indem Sie grundlegende Array-Algorithmen von Grund auf implementieren.

Sie arbeiten mit Arrays von Storyfragment-Objekten. Jedes Fragmentobjekt hat die folgenden Eigenschaften:

| Eigenschaft | Beschreibung | Beispielwert |
|-|-|-|
| `id` | Eine positive ganze Zahl, die die Position des Fragments in der Geschichte angibt | `3` |
| `text` | Der eigentliche Inhalt der Geschichte | `"and I use Arch btw.\""` |

In diesem Labor erhalten Sie ein vorbefülltes Array namens `shuffledFragments`.

**Ziel**: Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Labor abzuschließen.

**User Stories**

1. Sie sollten das vorbefüllte Array `shuffledFragments` nicht verändern.
2. Sie sollten eine Funktion namens `compactFragments` erstellen, die ein Array von Fragmenten entgegennimmt und ein neues Array zurückgibt, in dem alle undefined-Elemente entfernt sind. Wenn die Funktion undefined-Elemente entfernt, sollte sie eine Nachricht in der Konsole ausgeben. Die Nachricht sollte mit dem Präfix `[COMPACTED]` beginnen.
3. Sie sollten eine Variable namens `compactedShuffledFragments` deklarieren und ihr das Ergebnis des Aufrufs von `compactFragments` mit dem Array `shuffledFragments` zuweisen.
4. Sie sollten eine Funktion namens `sortFragments` erstellen, die ein Array von Fragmenten ohne `undefined`-Elemente entgegennimmt und ein neues Array zurückgibt, das nach der Eigenschaft `id` aufsteigend sortiert ist, wobei Fragmente mit derselben `id` in ihrer ursprünglichen Reihenfolge bleiben. Sie sollten nicht die eingebaute `sort`-Methode von JavaScript verwenden.
5. Sie sollten eine Variable namens `sortedFragments` deklarieren und ihr das Ergebnis des Aufrufs von `sortFragments` mit dem Array `compactedShuffledFragments` zuweisen.
6. Sie sollten eine Funktion namens `dedupeFragments` erstellen, die ein sortiertes Array von Fragmenten entgegennimmt und ein neues Array zurückgibt, in dem Duplikate entfernt sind und nur das erste Vorkommen erhalten bleibt. Sie sollten Duplikate als zwei oder mehr Fragmente definieren, die dieselbe `id` teilen. Für jede `id`, die dedupliziert wird, sollte die Funktion eine Nachricht in der Konsole ausgeben. Die Nachricht sollte mit dem Präfix `[DEDUPED]` beginnen.
7. Sie sollten eine Variable namens `dedupedFragments` deklarieren und ihr das Ergebnis des Aufrufs von `dedupeFragments` mit dem Array `sortedFragments` zuweisen.
8. Sie sollten eine Funktion namens `fillMissingFragments` erstellen, die ein sortiertes Array von Fragmenten entgegennimmt und ein neues Array zurückgibt, in dem fehlende Fragmente mit Platzhalterobjekten aufgefüllt sind. Sie sollten fehlende Fragmente als Lücken in der Folge zwischen der niedrigsten und der höchsten `id` definieren. Die Platzhalterobjekte sollten das Format `{ id: missingId, text: "[...]" }` haben. Für jeden hinzugefügten Platzhalter sollte die Funktion eine Nachricht in der Konsole ausgeben. Die Nachricht sollte mit dem Präfix `[FILLED]` beginnen.
9. Sie sollten eine Variable namens `filledFragments` deklarieren und ihr das Ergebnis des Aufrufs von `fillMissingFragments` mit dem Array `dedupedFragments` zuweisen.
10. Sie sollten eine Funktion namens `assembleStory` erstellen, die ein sortiertes Array von Fragmenten entgegennimmt und einen einzelnen String zurückgibt, der alle Fragmenttexte enthält, getrennt durch Zeilenumbrüche.
11. Sie sollten `assembleStory` mit Ihren `filledFragments` verwenden, um die vollständige Geschichte in der Konsole anzuzeigen.
12. Ihre Funktionen `compactFragments`, `sortFragments`, `dedupeFragments`, `fillMissingFragments` und `assembleStory` sollten das Array, mit dem sie aufgerufen werden, nicht verändern.

**Beispiel**

Hier ist ein Beispiel für ein Array, das Storyfragmente enthält:

```js
const exampleArray = [
  { id: 3, text: "and I use Arch btw.\"" },
  ,
  { id: 1, text: "Naomi said:" },
  { id: 3, text: "and I use Arch btw.\"" },
];
```

Nach der Wiederherstellung der Geschichte aus `exampleArray` würde sie so aussehen:

```md
Naomi said:
[...]
and I use Arch btw."
```

# --hints--

Sie dürfen das vorbefüllte Array `shuffledFragments` nicht verändern.

```js
const shuffledFragments__unchanged = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];

assert.deepEqual(shuffledFragments, shuffledFragments__unchanged);
```

Sie sollten eine Funktion `compactFragments` haben.

```js
assert.isFunction(compactFragments);
```

`compactFragments([{ id: 1, text: "Hello" }, undefined])` sollte eine Nachricht ausgeben, die mit `[COMPACTED]` beginnt.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 1, text: "Hello" }, undefined]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

Wenn die Funktion `compactFragments` mit einem Array aufgerufen wird, das `undefined`-Elemente enthält, sollte sie eine Konsolennachricht ausgeben, die mit `[COMPACTED]` beginnt.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 10, text: "foo" }, undefined, { id: 11, text: "bar" }, undefined]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

`compactFragments([{ id: 1, text: "Hello" },,])` sollte eine Nachricht ausgeben, die mit `[COMPACTED]` beginnt.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 1, text: "Hello" },,]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

Wenn die Funktion `compactFragments` mit einem Array aufgerufen wird, das leere Stellen enthält, sollte sie eine Konsolennachricht ausgeben, die mit `[COMPACTED]` beginnt.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 10, text: "foo" },,{ id: 11, text: "bar" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

Wenn die Funktion `compactFragments` mit dem Array `[{ id: 1, text: "Hello" }, { id: 2, text: "World" }]` aufgerufen wird, sollte sie keine Konsolennachrichten ausgeben, die mit `[COMPACTED]` beginnen.

```js
const shuffledFragments__withoutUndefined = [{ id: 1, text: "Hello" }, { id: 2, text: "World" }];
const spy = __helpers.spyOn(console, 'log');
compactFragments(shuffledFragments__withoutUndefined);
assert.isFalse(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

`compactFragments([{ id: 1, text: "Hello" }, undefined, { id: 2, text: "World" }])` sollte `[{ id: 1, text: "Hello" }, { id: 2, text: "World" }]` zurückgeben.

```js
assert.deepEqual(
  compactFragments([{ id: 1, text: "Hello" }, undefined, { id: 2, text: "World" }]),
  [{ id: 1, text: "Hello" }, { id: 2, text: "World" }]
);
```

`compactFragments` sollte das übergebene Array nicht verändern.

```js
const input = [{ id: 1, text: "a" }, undefined, { id: 2, text: "b" }];
const copy = structuredClone(input);
compactFragments(input);
assert.deepEqual(input, copy);
```

`compactedShuffledFragments` sollte ein Array ohne `undefined`-Elemente sein.

```js
assert.isArray(compactedShuffledFragments);
assert.isFalse(compactedShuffledFragments.some(f => f === undefined));
```

Sie sollten eine Funktion `sortFragments` haben.

```js
assert.isFunction(sortFragments);
```

`sortFragments([{ id: 30, text: "c" }, { id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }])` sollte `[{ id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }, { id: 30, text: "c" }]` zurückgeben.

```js
assert.deepEqual(
  sortFragments([{ id: 30, text: "c" }, { id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }]),
  [{ id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }, { id: 30, text: "c" }]
);
```

`sortFragments` sollte die eingebaute `.sort()`-Methode nicht verwenden.

```js
function isBuiltInSortUsed() {
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => { sortUsed = true; };
  try {
    sortFragments([{ id: 2, text: "b" }, { id: 1, text: "a" }]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

`sortFragments` sollte das übergebene Array nicht verändern.

```js
const input = [{ id: 3, text: "c" }, { id: 1, text: "a" }];
const copy = structuredClone(input);
sortFragments(input);
assert.deepEqual(input, copy);
```

`sortedFragments` sollte für jedes Element eine `id` haben, die kleiner oder gleich der `id` des nächsten Elements ist.

```js
for (let i = 0; i < sortedFragments.length - 1; i++) {
  assert.isAtMost(sortedFragments[i].id, sortedFragments[i + 1].id);
}
```

Sie sollten eine Funktion `dedupeFragments` haben.

```js
assert.isFunction(dedupeFragments);
```

`dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }])` sollte `[{ id: 1, text: "first" }, { id: 2, text: "second" }]` zurückgeben.

```js
assert.deepEqual(
  dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }]),
  [{ id: 1, text: "first" }, { id: 2, text: "second" }]
);
```

`dedupeFragments([{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }, { id: 2, text: "dup" }, { id: 3, text: "c" }])` sollte für jede duplizierte `id` eine Nachricht mit dem Präfix `[DEDUPED]` ausgeben.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }, { id: 2, text: "dup" }, { id: 3, text: "c" }]);
const dedupedLogs = spy.calls.filter(call => call[0].match(/^\[DEDUPED\]/));
assert.equal(dedupedLogs.length, 2);
spy.restore();
```

`dedupeFragments([{ id: 1, text: "a" }, { id: 2, text: "b" }, { id: 2, text: "dup" }])` sollte eine Nachricht mit dem Präfix `[DEDUPED]` ausgeben, wenn die letzte `id` dupliziert ist.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "a" }, { id: 2, text: "b" }, { id: 2, text: "dup" }]);
const dedupedLogs = spy.calls.filter(call => call[0].match(/^\[DEDUPED\]/));
assert.equal(dedupedLogs.length, 1);
spy.restore();
```

Wenn `dedupeFragments` mit einem Array aufgerufen wird, das Duplikate enthält, sollte es für jede duplizierte `id` eine Nachricht ausgeben, die mit `[DEDUPED]` beginnt.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[DEDUPED\]/)));
spy.restore();
```

`dedupeFragments` sollte das übergebene Array nicht verändern.

```js
const input = [{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }];
const copy = structuredClone(input);
dedupeFragments(input);
assert.deepEqual(input, copy);
```

`dedupedFragments` sollte keine doppelten `id`-Werte enthalten.

```js
const ids = dedupedFragments.map(f => f.id);
assert.strictEqual(ids.length, new Set(ids).size);
```

Sie sollten eine Funktion `fillMissingFragments` haben.

```js
assert.isFunction(fillMissingFragments);
```

`fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }])` sollte `[{ id: 1, text: "a" }, { id: 2, text: "[...]" }, { id: 3, text: "c" }]` zurückgeben.

```js
assert.deepEqual(
  fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }]),
  [{ id: 1, text: "a" }, { id: 2, text: "[...]" }, { id: 3, text: "c" }]
);
```

`fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }])` sollte eine Nachricht ausgeben, die mit `[FILLED]` beginnt.

```js
const spy = __helpers.spyOn(console, 'log');
fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[FILLED\]/)));
spy.restore();
```

Wenn `fillMissingFragments` einen Platzhalter hinzufügt, sollte es eine Nachricht ausgeben, die mit `[FILLED]` beginnt.

```js
const spy = __helpers.spyOn(console, 'log');
fillMissingFragments([{ id: 5, text: "e" }, { id: 8, text: "h" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[FILLED\]/)));
spy.restore();
```

`fillMissingFragments` sollte das übergebene Array nicht verändern.

```js
const input = [{ id: 1, text: "a" }, { id: 3, text: "c" }];
const copy = structuredClone(input);
fillMissingFragments(input);
assert.deepEqual(input, copy);
```

`filledFragments` sollte keine Lücken in der `id`-Folge haben.

```js
for (let i = 0; i < filledFragments.length - 1; i++) {
  assert.strictEqual(filledFragments[i + 1].id, filledFragments[i].id + 1);
}
```

Sie sollten eine Funktion `assembleStory` haben.

```js
assert.isFunction(assembleStory);
```

`assembleStory([{ id: 1, text: "Hello" }, { id: 2, text: "World" }])` sollte den String `"Hello\nWorld"` zurückgeben.

```js
assert.strictEqual(
  assembleStory([{ id: 1, text: "Hello" }, { id: 2, text: "World" }]),
  "Hello\nWorld"
);
```

`assembleStory` sollte das übergebene Array nicht verändern.

```js
const input = [{ id: 1, text: "Hello" }, { id: 2, text: "World" }];
const copy = structuredClone(input);
assembleStory(input);
assert.deepEqual(input, copy);
```

Sie sollten `console.log` mit dem Ergebnis von `assembleStory(filledFragments)` aufrufen.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(.*assembleStory\s*\(\s*filledFragments\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const shuffledFragments = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];
```

# --solutions--

```js
const shuffledFragments = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];

const compactFragments = (fragments) => {
  let compacted = [];
  let numberOfUndefined = 0;

  for (const fragment of fragments) {
    if (fragment === undefined) {
      numberOfUndefined++;
    } else {
      compacted.push(fragment);
    }
  }

  if (numberOfUndefined > 0) {
    console.log(`[COMPACTED] Removed ${numberOfUndefined} undefined fragment(s).`);
  }

  return compacted;
};

const compactedShuffledFragments = compactFragments(shuffledFragments);

const sortFragments = (fragments) => {
  let sorted = [...fragments];

  for (let i = 0; i < sorted.length; i++) {
    for (let j = i; j > 0; j--) {
      if (sorted[j - 1].id > sorted[j].id) {
        swap(sorted, j - 1, j);
      }
    }
  }

  function swap(array, index1, index2) {
    [sorted[index1], sorted[index2]] = [sorted[index2], sorted[index1]];
  }

  return sorted;
};

const sortedFragments = sortFragments(compactedShuffledFragments);

const dedupeFragments = (sortedFragments) => {
  const shallowCopy = [...sortedFragments];
  let deduped = [shallowCopy.shift()];
  let duplicates = [];
  let currentIdOccurrence = 1;

  for (const fragmentToInsert of shallowCopy) {
    const previousFragment = deduped.at(-1);
    if (fragmentToInsert.id === previousFragment.id) {
      currentIdOccurrence++;
    } else {
      deduped.push(fragmentToInsert);

      if (currentIdOccurrence > 1) {
        duplicates.push({ id: previousFragment.id, occurrence: currentIdOccurrence });
      }

      currentIdOccurrence = 1;
    }
  }

  if (currentIdOccurrence > 1) {
    duplicates.push({ id: deduped.at(-1).id, occurrence: currentIdOccurrence });
  }

  for (const duplicate of duplicates) {
    console.log(`[DEDUPED] Fragment #${duplicate.id} appeared ${duplicate.occurrence} times, kept first occurrence.`);
  }

  return deduped;
};

const dedupedFragments = dedupeFragments(sortedFragments);

const fillMissingFragments = (sortedFragments) => {
  const shallowCopy = [...sortedFragments];
  let filled = [shallowCopy.shift()];

  for (const fragmentToInsert of shallowCopy) {
    const previousFragment = filled.at(-1);
    const numberOfMissingIDs = fragmentToInsert.id - 1 - previousFragment.id;
    for (let i = 1; i <= numberOfMissingIDs; i++) {
      const missingId = previousFragment.id + i;
      filled.push({ id: missingId, text: "[...]" });
      console.log(`[FILLED] Fragment #${missingId} missing, placeholder inserted.`)
    }
    filled.push(fragmentToInsert);
  }

  return filled;
};

const filledFragments = fillMissingFragments(dedupedFragments);

const assembleStory = (sortedFragments) => {
  const texts = [];
  for (const fragment of sortedFragments) {
    texts.push(fragment.text);
  }
  return texts.join("\n");
};

console.log("\nHere is the assembled story:\n");
console.log(assembleStory(filledFragments));
```
