---
id: 646ceb843412c74edee27a79
title: ステップ 22
challengeType: 0
dashedName: step-22
---

# --description--

これから三角形を描くためのCSSのテクニックを学びます。

これは猫の耳を描くのに使います。

クラスセレクターを使って、`.cat-right-ear` 要素に`height` と `width` のプロパティを`100px` に設定してください。`background-color` は`white` に設定します。左右のボーダーは`35px solid blue` に設定してください。上下のボーダーは`35px solid red` に設定します。

# --hints--

`.cat-right-ear` セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

`.cat-right-ear` セレクターには`height` プロパティを`100px` に設定してください。セミコロンを忘れないでください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height, '100px')
```

`.cat-right-ear` セレクターには`width` プロパティを`100px` に設定してください。セミコロンを忘れないでください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width, '100px')
```

`.cat-right-ear` セレクターには`background-color` プロパティを`white` に設定してください。セミコロンを忘れないでください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor, 'white')
```

`.cat-right-ear` セレクターには`border-left` プロパティを`35px solid blue` に設定してください。セミコロンを忘れないでください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid blue')
```

`.cat-right-ear` セレクターには`border-right` プロパティを`35px solid blue` に設定してください。セミコロンを忘れないでください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid blue')
```

`.cat-right-ear` セレクターには`border-top` プロパティを`35px solid red` に設定してください。セミコロンを忘れないでください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop, '35px solid red')
```

`.cat-right-ear` セレクターには`border-bottom` プロパティを`35px solid red` に設定してください。セミコロンを忘れないでください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom, '35px solid red')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
