---
id: 646c5ffef5598d449b52ec12
title: ステップ 19
challengeType: 0
dashedName: step-19
---

# --description--

次は猫の耳に取り組みましょう。右耳と左耳があり、それぞれの中に内耳があります。

`.cat-head` 要素の中に、クラス `div` を持つ `cat-ears` 要素を作成してください。

# --hints--

既存のクラス `div` を持つ `cat-head` 要素は変更しないでください。

```js
assert(document.querySelectorAll('div.cat-head').length === 1);
```

`div` 要素の中に `.cat-head` 要素を1つ作成してください。

```js
assert(document.querySelectorAll('.cat-head div').length === 1);
```

`div` 要素にはクラス `cat-ears` を付けてください。

```js
assert(document.querySelectorAll('.cat-head div')[0]?.getAttribute('class') === 'cat-ears');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      --fcc-editable-region--
      <div class="cat-head">
        
      </div>
      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
