---
id: 646c47867800472a4ed5d2ea
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

基本的なHTML構造から始めましょう。`DOCTYPE`の`html`参照と、`html`属性が`lang`に設定された`en`要素を追加してください。また、`head`要素の中に`body`と`html`要素も追加してください。

# --hints--

`DOCTYPE`の`html`宣言があるはずです。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

`DOCTYPE`宣言はHTMLの最初にあるべきです。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`html`の`lang`属性を持つ開く`en`タグがあるはずです。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

閉じる`html`タグがあるはずです。

```js
assert(code.match(/<\/html>/i));
```

開く`head`タグがあるはずです。

```js
assert(code.match(/<head>/i));
```

閉じる`head`タグがあるはずです。

```js
assert(code.match(/<\/head>/i));
```

開く`body`タグがあるはずです。

```js
assert(code.match(/<body>/i));
```

閉じる`body`タグがあるはずです。

```js
assert(code.match(/<\/body>/i));
```

`body`要素の後に`head`要素が来るはずです。

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

`head`と`body`要素は`html`要素の中にあるはずです。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
