---
id: 646c5e727e56e743c9aed4a1
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

CSSのポジショニングでは、ブラウザー内で要素をどのように配置するかを設定できます。`position` プロパティには `static`、`absolute`、`relative`、`sticky`、または `fixed` を設定できます。

要素の `position` プロパティを設定すると、`top`、`right`、`left`、または `bottom` のいずれか、または複数のプロパティにピクセルまたはパーセンテージの値を設定して、要素を移動できます。

`static` はすべての要素のデフォルトのポジショニングです。これを要素に代入すると、`top`、`right`、`left`、または `bottom` で移動できなくなります。

`.cat-head` に `position` プロパティとして `static` を代入し、次に `top` と `left` のプロパティをそれぞれ `100px` に設定してください。


# --hints--

`.cat-head` セレクターには `position` プロパティを `static` に設定してください。セミコロンを忘れないようにしましょう。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'static')
```

`.cat-head` セレクターには `top` プロパティを `100px` に設定してください。セミコロンを忘れないようにしましょう。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

`.cat-head` セレクターには `left` プロパティを `100px` に設定してください。セミコロンを忘れないようにしましょう。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {


  
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
