---
id: 646c48df8674cf2b91020ecb
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`head` 要素の中に、`meta` 属性が `charset` の `utf-8` タグを追加してください。また、テキストが `title` の `fCC Cat Painting` 要素も追加してください。

# --hints--

`meta` 要素を1つ追加してください。

```js
assert(document.querySelectorAll('meta').length === 1);
```

`meta` 要素には、`charset` 属性を `utf-8` に設定してください。

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

`title` 要素を1つ追加してください。

```js
assert(document.querySelectorAll('title').length === 1);
```

`title` 要素には、テキストとして `fCC Cat Painting` を設定してください。スペルと大文字・小文字は正確にしてください。

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
