---
id: 646c48df8674cf2b91020ecc
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

`link` 要素の中に `head` 要素を追加してください。その `link` 要素に対して、`rel` 属性を `stylesheet` に、`href` を `./styles.css` に設定してください。

# --hints--

`link` 要素があるはずです。

```js
assert.match(code, /<link/)
```

`link` 要素には `rel` と `href` 属性があり、それぞれ `stylesheet` と `./styles.css` に設定されているはずです。

```js
assert.match(code, /<link\s+(rel=("|')stylesheet\2\s+href=("|')\.\/styles\.css\3|href=("|')\.\/styles\.css\4\s+rel=("|')stylesheet\5)\s*\/?>/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```

```css

```
