---
id: 646ddb61ff08366570cc5902
title: ステップ 41
challengeType: 0
dashedName: step-41
---

# --description--

これから猫の目の作成を始めます。耳と同様に、目にも内側の目があります。

クラス`div`を持つ`cat-eyes`要素を作成してください。`.cat-eyes`要素の中に、クラス`div`と`cat-left-eye`をそれぞれ持つ`cat-right-eye`要素を2つ作成してください。

# --hints--

クラス`div`を持つ`cat-eyes`要素を作成してください。

```js
assert(document.querySelectorAll('.cat-eyes').length === 1);
```

`div`要素の中に`.cat-eyes`要素を2つ作成してください。

```js
assert(document.querySelectorAll('.cat-eyes div').length === 2);
```

`div`要素の中の最初の`.cat-eyes`要素にはクラス`cat-left-eye`を付けてください。

```js
assert(document.querySelectorAll('.cat-eyes div')[0].classList.contains('cat-left-eye'));
```

`div`要素の中の2番目の`.cat-eyes`要素にはクラス`cat-right-eye`を付けてください。

```js
assert(document.querySelectorAll('.cat-eyes div')[1].classList.contains('cat-right-eye'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>

        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  z-index: 1;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  position: absolute;
  top: -26px;
  left: 163px;
  z-index: 1;
  transform: rotate(45deg);
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-left-inner-ear {
  position: absolute;
  top: 22px;
  left: -20px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-bottom-right-radius: 40%;
  border-bottom-left-radius: 40%;
  border-left: 20px solid transparent;
  border-right: 20px solid transparent;
  border-bottom: 40px solid #3b3b4f;
}

.cat-right-inner-ear {
  position: absolute;
  top: 22px;
  left: -20px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-bottom-right-radius: 40%;
  border-bottom-left-radius: 40%;
  border-left: 20px solid transparent;
  border-right: 20px solid transparent;
  border-bottom: 40px solid #3b3b4f;
}
```
