---
id: 646dd556d524bc61c0139bd6
title: ステップ 35
challengeType: 0
dashedName: step-35
---

# --description--

ほとんどの猫は、耳の外側と内側で色が異なります。あなたも同じように実装できます。だからこそ、少し前に右耳と左耳の内側それぞれに `div` 要素を定義しました。

クラスセレクターを使って、`.cat-left-inner-ear` 要素に左右のボーダーをそれぞれ `20px solid transparent` で設定してください。また、下のボーダーには `40px solid #3b3b4f` を設定してください。

# --hints--

`.cat-left-inner-ear` セレクターがあるはずです。

```js 
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-inner-ear'))
```

`.cat-left-inner-ear` セレクターには `border-left` プロパティを `20px solid transparent` に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-inner-ear')?.borderLeft === '20px solid transparent')
```

`.cat-left-inner-ear` セレクターには `border-right` プロパティを `20px solid transparent` に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-inner-ear')?.borderRight === '20px solid transparent')
```

`.cat-left-inner-ear` セレクターには `border-bottom` プロパティを `40px solid #3b3b4f` に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-inner-ear')?.borderBottom === '40px solid rgb(59, 59, 79)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  z-index: 1;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  position: absolute;
  top: -26px;
  left: 163px;
  z-index: 1;
  transform: rotate(45deg);
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

--fcc-editable-region--

--fcc-editable-region--
```
