---
id: 6476f7a4827bcc61682f2347
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

次の position プロパティは `absolute` です。`absolute` プロパティに `position` の値を使うと、要素は文書の通常のフローから外れ、その位置は `top`、`right`、`bottom`、`left` のプロパティによって決まります。

`position` 要素の `.cat-head` プロパティを `absolute` に設定し、次に `top` と `left` のプロパティを任意の正のピクセル値に設定してください。

<!-- **注意**: ここで `top`、`left`、`bottom`、`right` のプロパティを試すことができますが、テストは `top` の `300px` と `400px` の left のみで合格します。 -->

# --hints--

`.cat-head` セレクターには `position` プロパティを `absolute` に設定してください。セミコロンを忘れないでください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

`.cat-head` セレクターには `top` プロパティを `px` の任意の数値に設定してください。セミコロンを忘れないでください。

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/)
```

`.cat-head` セレクターには `left` プロパティを `px` の任意の数値に設定してください。セミコロンを忘れないでください。

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
