---
id: 64770351e8586671ec0911f0
title: ステップ 15
challengeType: 0
dashedName: step-15
---

# --description--

最後の position プロパティの値は `sticky` です。`sticky` の位置指定は `relative` と `fixed` のハイブリッドです。スクロール位置に基づいて、要素をその親要素やビューポート内の特定の位置に**固定**できます。

`position` の `.cat-head` プロパティの値を `sticky` に変更し、`top` を `0` に設定してから、`left` とその値を削除してください。

**注意**：`sticky` の動作を確認するには、`.cat-head` `div` 要素の前後にいくつかテキストを配置する必要があります。その後スクロールすると、`.cat-head` が上部に固定されて動かなくなるのが見えます。

# --hints--

`.cat-head` セレクターには `position` プロパティを `sticky` に設定してください。セミコロンを忘れずに追加しましょう。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'sticky')
```

`.cat-head` セレクターには `top` プロパティを `0` に設定してください。セミコロンを忘れずに追加しましょう。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

コード内に `left` プロパティとその値が含まれていないことを確認してください。

```js
assert.notMatch(code, /left:\s*100px;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
