---
id: 647d855ac12cd436059acd39
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

ユニバーサルセレクターを使って、CSSに`box-sizing: border-box;`を追加してください。これにより、要素の指定された幅と高さにパディングとボーダーが含まれるようになります。

# --hints--

ユニバーサルセレクター（`*`）があるはずです。

```js
// assert.match(code, /\*\s*{\s*.*\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*'))
```

ユニバーサルセレクターには`box-sizing`プロパティが`border-box`に設定されている必要があります。

```js
// assert.match(code, /\*\s*{\s*box-sizing:\s*border-box;?\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
