---
id: 6676a8b01e56ec1a1e07c202
title: ステップ 23
challengeType: 0
dashedName: step-23
---

# --description--

白い四角形に太いボーダーがついているのがわかります。ボーダーは斜めのエッジになっていて、三角形を作るのに使えます。

同じクラスセレクター内で `height` と `width` を `0` に調整してください。左、右、上のボーダーを `transparent` に設定します。

`background-color` プロパティを削除すると、三角形が見えるはずです。

# --hints--


`.cat-right-ear` セレクターがあるはずです。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

`.cat-right-ear` セレクターには `height` プロパティが `0` に設定されているはずです。セミコロンを忘れないでください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height == '0px')
```

`.cat-right-ear` セレクターには `width` プロパティが `0` に設定されているはずです。セミコロンを忘れないでください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width === '0px')
```

`.cat-right-ear` セレクターには `border-left` プロパティが `35px solid transparent` に設定されているはずです。セミコロンを忘れないでください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft === '35px solid transparent')
```

`.cat-right-ear` セレクターには `border-right` プロパティが `35px solid transparent` に設定されているはずです。セミコロンを忘れないでください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight === '35px solid transparent')
```

`.cat-right-ear` セレクターには `border-top` プロパティが `35px solid transparent` に設定されているはずです。セミコロンを忘れないでください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop === '35px solid transparent')
```


`.cat-right-ear` セレクターには `border-bottom` プロパティが `35px solid red` に設定されているはずです。セミコロンを忘れないでください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom === '35px solid red')
```

`.cat-right-ear` セレクターには `background-color` プロパティがあってはいけません。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 100px;
  width: 100px;
  background-color: white;
  border-left: 35px solid blue;
  border-right: 35px solid blue;
  border-top: 35px solid red;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
