---
id: 5a24c314108439a4d4036156
title: 非同期アクションを処理するためにミドルウェアを使う
challengeType: 6
forumTopicId: 301451
dashedName: use-middleware-to-handle-asynchronous-actions
---

# --description--

これまでのチャレンジでは非同期アクションについて触れていませんでしたが、非同期アクションはウェブ開発において避けられないものです。Reduxアプリで非同期のエンドポイントを呼び出す必要が出てくるので、こうしたリクエストをどう処理するかが問題になります。Reduxにはこの目的のために設計されたミドルウェア、Redux Thunkミドルウェアがあります。ここではReduxでこれを使う方法を簡単に説明します。

Redux Thunkミドルウェアを含めるには、`Redux.applyMiddleware()`に引数として渡します。この文は`createStore()`関数の2番目のオプションのパラメータとして渡されます。エディターの下部にあるコードを見てください。非同期アクションを作成するには、アクションクリエーター内で`dispatch`を引数に取る関数を返します。この関数内でアクションをdispatchし、非同期リクエストを実行できます。

この例では、非同期リクエストは`setTimeout()`呼び出しでシミュレートされています。非同期処理を開始する前にアクションをdispatchするのが一般的で、これによりアプリケーションの状態がデータをリクエスト中であることを認識できます（例えばローディングアイコンを表示するなど）。そしてデータを受け取ったら、データをペイロードとして持ち、処理が完了したことを示す別のアクションをdispatchします。

この特別なアクションクリエーターに`dispatch`をパラメータとして渡していることを覚えておいてください。これを使ってアクションをdispatchし、アクションを直接dispatchに渡せばミドルウェアが残りの処理を行います。

# --instructions--

`handleAsync()`アクションクリエーター内で2つのdispatchを記述してください。`requestingData()`（シミュレートされたAPI呼び出し）の前に`setTimeout()`をdispatchします。そして（仮の）データを受け取った後に、このデータを渡して`receivedData()`アクションをdispatchしてください。これでReduxで非同期アクションを処理する方法がわかりました。その他はこれまで通りに動作します。

# --hints--

`requestingData`アクションクリエーターは、`REQUESTING_DATA`の値と等しいtypeのオブジェクトを返す必要があります。

```js
assert(requestingData().type === REQUESTING_DATA);
```

`receivedData`アクションクリエーターは、`RECEIVED_DATA`の値と等しいtypeのオブジェクトを返す必要があります。

```js
assert(receivedData('data').type === RECEIVED_DATA);
```

`asyncDataReducer`は関数である必要があります。

```js
assert(typeof asyncDataReducer === 'function');
```

`requestingData`アクションクリエーターをdispatchすると、ストアの`state`プロパティのfetchingが`true`に更新されます。

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(requestingData());
    const reqState = store.getState();
    return initialState.fetching === false && reqState.fetching === true;
  })()
);
```

`handleAsync`をdispatchすると、データリクエストアクションをdispatchし、その後遅延の後に受信データアクションをdispatchします。

```js
assert(
  (function () {
    const noWhiteSpace = __helpers.removeWhiteSpace(handleAsync.toString());
    return (
      noWhiteSpace.includes('dispatch(requestingData())') === true &&
      noWhiteSpace.includes('dispatch(receivedData(data))') === true
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
const REQUESTING_DATA = 'REQUESTING_DATA'
const RECEIVED_DATA = 'RECEIVED_DATA'

const requestingData = () => { return {type: REQUESTING_DATA} }
const receivedData = (data) => { return {type: RECEIVED_DATA, users: data.users} }

const handleAsync = () => {
  return function(dispatch) {
    // Dispatch request action here

    setTimeout(function() {
      let data = {
        users: ['Jeff', 'William', 'Alice']
      }
      // Dispatch received data action here

    }, 2500);
  }
};

const defaultState = {
  fetching: false,
  users: []
};

const asyncDataReducer = (state = defaultState, action) => {
  switch(action.type) {
    case REQUESTING_DATA:
      return {
        fetching: true,
        users: []
      }
    case RECEIVED_DATA:
      return {
        fetching: false,
        users: action.users
      }
    default:
      return state;
  }
};

const store = Redux.createStore(
  asyncDataReducer,
  Redux.applyMiddleware(ReduxThunk.default)
);
```

# --solutions--

```js
const REQUESTING_DATA = 'REQUESTING_DATA'
const RECEIVED_DATA = 'RECEIVED_DATA'

const requestingData = () => { return {type: REQUESTING_DATA} }
const receivedData = (data) => { return {type: RECEIVED_DATA, users: data.users} }

const handleAsync = () => {
  return function(dispatch) {
    dispatch(requestingData());
    setTimeout(function() {
      let data = {
        users: ['Jeff', 'William', 'Alice']
      }
      dispatch(receivedData(data));
    }, 2500);
  }
};

const defaultState = {
  fetching: false,
  users: []
};

const asyncDataReducer = (state = defaultState, action) => {
  switch(action.type) {
    case REQUESTING_DATA:
      return {
        fetching: true,
        users: []
      }
    case RECEIVED_DATA:
      return {
        fetching: false,
        users: action.users
      }
    default:
      return state;
  }
};

const store = Redux.createStore(
  asyncDataReducer,
  Redux.applyMiddleware(ReduxThunk.default)
);
```
