---
id: 5a24c314108439a4d4036153
title: ストアリスナーを登録する
challengeType: 6
forumTopicId: 301446
dashedName: register-a-store-listener
---

# --description--

Redux `store` オブジェクトでアクセスできるもう一つのメソッドは `store.subscribe()` です。これはストアにリスナー関数を登録でき、ストアに対してアクションがディスパッチされるたびに呼び出されます。このメソッドの簡単な使い方としては、アクションを受け取ってストアが更新されるたびにメッセージをログに出力する関数をストアに登録することです。

# --instructions--

ストアがアクションを受け取るたびにグローバル変数 `count` をインクリメントするコールバック関数を書き、この関数を `store.subscribe()` メソッドに渡してください。`store.dispatch()` が連続して3回呼び出され、それぞれ直接アクションオブジェクトを渡しているのがわかります。アクションのディスパッチ間のコンソール出力を見て、更新が行われていることを確認してください。

# --hints--

ストアで `ADD` アクションをディスパッチすると、状態が `1` だけインクリメントされるはずです。

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch({ type: 'ADD' });
    const newState = store.getState();
    return newState === initialState + 1;
  })()
);
```

`store.subscribe` を使ってストアにリスナー関数が登録されている必要があります。

```js
assert.match(code, /store\s*\.\s*subscribe\(/gm);
```

`store.subscribe` は関数を受け取るべきです。

```js
assert.match(code, /(\s*function\s*)|(\s*\(\s*\)\s*=>)/gm);
```

`store.subscribe` に渡された関数は呼び出されてはいけません。

```js
assert.notMatch(code, /store\.subscribe\(.+\(\)\)/);
```

`store.subscribe` へのコールバックは、ストアが更新されるたびにグローバル `count` 変数もインクリメントする必要があります。

```js
assert.strictEqual(store.getState(), count);
```

# --seed--

## --seed-contents--

```js
const ADD = 'ADD';

const reducer = (state = 0, action) => {
  switch(action.type) {
    case ADD:
      return state + 1;
    default:
      return state;
  }
};

const store = Redux.createStore(reducer);

// Global count variable:
let count = 0;

// Change code below this line

// Change code above this line

store.dispatch({type: ADD});
console.log(count);
store.dispatch({type: ADD});
console.log(count);
store.dispatch({type: ADD});
console.log(count);
```

# --solutions--

```js
const ADD = 'ADD';

const reducer = (state = 0, action) => {
  switch(action.type) {
    case ADD:
      return state + 1;
    default:
      return state;
  }
};

const store = Redux.createStore(reducer);
 let count = 0;
// Change code below this line

store.subscribe( () =>
 {
 count++;
 }
);

// Change code above this line

store.dispatch({type: ADD});
store.dispatch({type: ADD});
store.dispatch({type: ADD});
```
