---
id: 5a24c314108439a4d403615a
title: 配列からアイテムを削除する
challengeType: 6
forumTopicId: 301447
dashedName: remove-an-item-from-an-array
---

# --description--

配列からアイテムを削除する練習をしましょう。ここでもスプレッド演算子を使うことができます。その他に便利なJavaScriptのメソッドには`slice()`や`concat()`があります。

# --instructions--

リデューサーとアクションクリエーターは、アイテムのインデックスに基づいて配列からアイテムを削除するように変更されています。特定のインデックスのアイテムが削除された新しい状態の配列を返すように、リデューサーの記述を完成させてください。

# --hints--

Reduxのストアは存在し、状態は`[0,1,2,3,4,5]`と等しく初期化されている必要があります。

```js
const initialState = store.getState();
assert.isArray(initialState) 
assert.deepEqual(initialState, [0, 1, 2, 3, 4, 5])
```

`removeItem`と`immutableReducer`はどちらも関数である必要があります。

```js
assert(
  typeof removeItem === 'function' && typeof immutableReducer === 'function'
);
```

`removeItem`アクションクリエーターをディスパッチすると、状態からアイテムが削除され、状態は変更されてはいけません。

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(removeItem(3));
const state_1 = store.getState();
store.dispatch(removeItem(2));
const state_2 = store.getState();
store.dispatch(removeItem(0));
store.dispatch(removeItem(0));
store.dispatch(removeItem(0));
const state_3 = store.getState();

assert(isFrozen)
assert.deepEqual(state_1, [0, 1, 2, 4, 5]) 
assert.deepEqual(state_2, [0, 1, 4, 5]) 
assert.deepEqual(state_3, [5])

```

# --seed--

## --seed-contents--

```js
const immutableReducer = (state = [0,1,2,3,4,5], action) => {
  switch(action.type) {
    case 'REMOVE_ITEM':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const removeItem = (index) => {
  return {
    type: 'REMOVE_ITEM',
    index
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const immutableReducer = (state = [0,1,2,3,4,5], action) => {
  switch(action.type) {
    case 'REMOVE_ITEM':
      return [
        ...state.slice(0, action.index),
        ...state.slice(action.index + 1)
      ];
    default:
      return state;
  }
};

const removeItem = (index) => {
  return {
    type: 'REMOVE_ITEM',
    index
  }
}

const store = Redux.createStore(immutableReducer);
```
