---
id: 5a24c314108439a4d4036152
title: アクションタイプに const を使う
challengeType: 6
forumTopicId: 301450
dashedName: use-const-for-action-types
---

# --description--

Redux を使う際の一般的な慣例として、アクションタイプを読み取り専用の定数として代入し、それらの定数を使用する場所で参照します。作業中のコードをリファクタリングして、アクションタイプを `const` 宣言として書き換えてください。

# --instructions--

`LOGIN` と `LOGOUT` を `const` 値として宣言し、それぞれ文字列 `'LOGIN'` と `'LOGOUT'` を代入してください。次に、`authReducer()` とアクションクリエーターを編集して、文字列の代わりにこれらの定数を参照するようにしてください。

**注意:** 定数は一般的にすべて大文字で書く慣例があり、これは Redux でも標準的な慣習です。

# --hints--

`loginUser` 関数を呼び出すと、`type` プロパティが文字列 `LOGIN` に設定されたオブジェクトが返されるはずです。

```js
assert(loginUser().type === 'LOGIN');
```

`logoutUser` 関数を呼び出すと、`type` プロパティが文字列 `LOGOUT` に設定されたオブジェクトが返されるはずです。

```js
assert(logoutUser().type === 'LOGOUT');
```

ストアは、`login` プロパティが `false` に設定されたオブジェクトで初期化されるべきです。

```js
assert(store.getState().authenticated === false);
```

`loginUser` をディスパッチすると、ストアの状態の `login` プロパティが `true` に更新されるはずです。

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(loginUser());
    const afterLogin = store.getState();
    return (
      initialState.authenticated === false && afterLogin.authenticated === true
    );
  })()
);
```

`logoutUser` をディスパッチすると、ストアの状態の `login` プロパティが `false` に更新されるはずです。

```js
assert(
  (function () {
    store.dispatch(loginUser());
    const loggedIn = store.getState();
    store.dispatch(logoutUser());
    const afterLogout = store.getState();
    return (
      loggedIn.authenticated === true && afterLogout.authenticated === false
    );
  })()
);
```

`authReducer` 関数は、switch文で複数のアクションタイプを処理するべきです。

```js
assert(
    (function () {
      return (
        typeof authReducer === 'function' &&
        code.toString().includes('switch') &&
        code.toString().includes('case') &&
        code.toString().includes('default')
      );
    })()
  );
```

`LOGIN` と `LOGOUT` は `const` 値として宣言され、文字列 `LOGIN` と `LOGOUT` が代入されるべきです。

```js
const noWhiteSpace = __helpers.removeWhiteSpace(code);
assert(LOGIN === 'LOGIN' && LOGOUT === 'LOGOUT')
assert(noWhiteSpace.includes('const'))
```

アクションクリエーターとリデューサーは、`LOGIN` と `LOGOUT` の定数を参照するべきです。

```js
  assert(
    (function () {
      const noWhiteSpace = __helpers.removeWhiteSpace(
        code.toString()
      );
      return (
        noWhiteSpace.includes('caseLOGIN:') &&
        noWhiteSpace.includes('caseLOGOUT:') &&
        noWhiteSpace.includes('type:LOGIN') &&
        noWhiteSpace.includes('type:LOGOUT')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```js


const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {
    case 'LOGIN': 
      return {
        authenticated: true
      }
    case 'LOGOUT': 
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```

# --solutions--

```js
const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {

    case LOGIN:
      return {
        authenticated: true
      }

    case LOGOUT:
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: LOGIN
  }
};

const logoutUser = () => {
  return {
    type: LOGOUT
  }
};
```
