---
id: 5a24c314108439a4d4036158
title: 状態を決して変更しない
challengeType: 6
forumTopicId: 301445
dashedName: never-mutate-state
---

# --description--

この最後のチャレンジでは、Reduxにおける状態の不変性という重要な原則を守るためのいくつかの方法を説明します。不変の状態とは、状態を直接変更せずに、新しい状態のコピーを返すことを意味します。

Reduxアプリの状態を時間経過でスナップショットとして撮ると、`state 1`、`state 2`、`state 3`、`state 4`、`...`のように、各状態は前の状態と似ているかもしれませんが、それぞれが異なるデータの塊であることがわかります。この不変性こそが、タイムトラベルデバッグのような機能を実現しているのです。

Reduxはストアやリデューサーで状態の不変性を強制しません。その責任はプログラマーにあります。幸いなことに、JavaScript（特にES6）は、状態が`string`、`number`、`array`、または`object`であっても、その不変性を守るために使える便利なツールをいくつか提供しています。文字列や数値はプリミティブ値であり、もともと不変です。つまり、3は常に3であり、数値3の値を変更することはできません。一方、`array`や`object`は変更可能です。実際には、状態は`array`や`object`で構成されることが多いでしょう。これらは多くの種類の情報を表すのに便利なデータ構造です。

# --instructions--

コードエディターには、ToDoアイテムを管理するための`store`と`reducer`があります。リデューサーの`ADD_TO_DO`ケースを完成させて、新しいToDoを状態に追加してください。標準のJavaScriptやES6でこれを実現する方法はいくつかあります。`action.todo`からのアイテムを末尾に追加した新しい配列を返す方法を探してみてください。

# --hints--

Reduxストアは存在し、コードエディター内の`todos`配列と同じ状態で初期化されているはずです。

```js
assert(
  (function () {
    const todos = [
      'Go to the store',
      'Clean the house',
      'Cook dinner',
      'Learn to code'
    ];
    const initialState = store.getState();
    return (
      Array.isArray(initialState) && initialState.join(',') === todos.join(',')
    );
  })()
);
```

`addToDo`と`immutableReducer`はどちらも関数である必要があります。

```js
assert(typeof addToDo === 'function' && typeof immutableReducer === 'function');
```

Reduxストアで`ADD_TO_DO`タイプのアクションをディスパッチすると、`todo`アイテムが追加され、状態は変更されてはいけません。

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(addToDo('__TEST__TO__DO__'));
const finalState = store.getState();
const expectedState = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
  '__TEST__TO__DO__'
];
assert(isFrozen);
assert.deepEqual(finalState, expectedState);
```

# --seed--

## --seed-contents--

```js
const ADD_TO_DO = 'ADD_TO_DO';

// A list of strings representing tasks to do:
const todos = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
];

const immutableReducer = (state = todos, action) => {
  switch(action.type) {
    case ADD_TO_DO:
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: ADD_TO_DO,
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const ADD_TO_DO = 'ADD_TO_DO';

const todos = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
];

const immutableReducer = (state = todos, action) => {
  switch(action.type) {
    case ADD_TO_DO:
      return state.concat(action.todo);
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: ADD_TO_DO,
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```
