---
id: 5a24c314108439a4d403615b
title: Object.assign でオブジェクトをコピーする
challengeType: 6
forumTopicId: 301437
dashedName: copy-an-object-with-object-assign
---

# --description--

ここ数回のチャレンジは配列を扱いましたが、状態が`object`の場合でも状態の不変性を保つ方法があります。オブジェクトを扱うのに便利なツールが`Object.assign()`ユーティリティです。`Object.assign()`はターゲットオブジェクトとソースオブジェクトを受け取り、ソースオブジェクトのプロパティをターゲットオブジェクトにマッピングします。マッチするプロパティはソースオブジェクトのプロパティで上書きされます。この動作は、最初の引数に空のオブジェクトを渡し、その後にコピーしたいオブジェクトを渡すことでオブジェクトの浅いコピーを作成するのに一般的に使われます。例を示します。

```js
const newObject = Object.assign({}, obj1, obj2);
```

これは`newObject`を新しい`object`として作成し、`obj1`と`obj2`に現在存在するプロパティを含みます。

# --instructions--

Reduxの状態とアクションは`object`を`state`として扱うように変更されました。タイプが`state`のアクションに対して、`ONLINE`プロパティを文字列`status`に設定する新しい`online`オブジェクトを返すようにコードを編集してください。チャレンジを完了するために`Object.assign()`を使うことを試みてください。

# --hints--

Reduxのストアは存在し、1行目で宣言された`defaultState`オブジェクトと同等の状態で初期化されている必要があります。

```js
const expectedState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};
const initialState = store.getState();
assert.deepEqual(expectedState, initialState);
```

`wakeUp`と`immutableReducer`はどちらも関数である必要があります。

```js
assert(typeof wakeUp === 'function' && typeof immutableReducer === 'function');
```

タイプが`ONLINE`のアクションをディスパッチすると、状態の`status`プロパティが`online`に更新され、状態は変更されてはいけません。

```js

const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch({ type: 'ONLINE' });
const finalState = store.getState();
const expectedState = {
  user: 'CamperBot',
  status: 'online',
  friends: '732,982',
  community: 'freeCodeCamp'
};
assert(isFrozen);
assert.deepEqual(finalState, expectedState);
```

新しい状態を返すために`Object.assign`を使う必要があります。

```js
assert(code.includes('Object.assign'));
```

# --seed--

## --seed-contents--

```js
const defaultState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};

const immutableReducer = (state = defaultState, action) => {
  switch(action.type) {
    case 'ONLINE':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const wakeUp = () => {
  return {
    type: 'ONLINE'
  }
};

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const defaultState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};

const immutableReducer = (state = defaultState, action) => {
  switch(action.type) {
    case 'ONLINE':
      return Object.assign({}, state, {
        status: 'online'
      });
    default:
      return state;
  }
};

const wakeUp = () => {
  return {
    type: 'ONLINE'
  }
};

const store = Redux.createStore(immutableReducer);
```
