---
id: 5a24c314108439a4d4036159
title: 配列にスプレッド演算子を使う
challengeType: 6
forumTopicId: 301452
dashedName: use-the-spread-operator-on-arrays
---

# --description--

Reduxで状態の不変性を保つためのES6の解決策の一つがスプレッド演算子です：`...`。スプレッド演算子はさまざまな用途がありますが、その一つは既存の配列から新しい配列を作成するという前のチャレンジに適しています。これは比較的新しいですが、よく使われる構文です。例えば、配列`myArray`があって次のように書くと：

```js
let newArray = [...myArray];
```

`newArray`は`myArray`のクローンになります。両方の配列はメモリ上で別々に存在しています。`newArray.push(5)`のようなミューテーションを行っても、`myArray`は変わりません。`...`は`myArray`の値を新しい配列に*展開*しています。配列をクローンしつつ新しい配列に追加の値を加えたい場合は、`[...myArray, 'new value']`のように書けます。これは`myArray`の値と文字列`new value`を最後の値として含む新しい配列を返します。スプレッド構文はこのように配列の構成で複数回使えますが、浅いコピーしか作らないことに注意してください。つまり、一次元配列に対してのみ不変な配列操作を提供します。

# --instructions--

スプレッド演算子を使って、to-doが追加されたときに状態の新しいコピーを返してください。

# --hints--

Reduxのストアは存在し、状態は`["Do not mutate state!"]`と等しい初期値で初期化されている必要があります。

```js
assert(
  (function () {
    const initialState = store.getState();
    return (
      Array.isArray(initialState) === true &&
      initialState[0] === 'Do not mutate state!'
    );
  })()
);
```

`addToDo`と`immutableReducer`はどちらも関数である必要があります。

```js
assert(typeof addToDo === 'function' && typeof immutableReducer === 'function');
```

Reduxストアに`ADD_TO_DO`タイプのアクションをディスパッチすると、`todo`の項目が追加され、状態はミューテーションされてはいけません。

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(addToDo('__TEST__TO__DO__'));
const finalState = store.getState();
const expectedState = ['Do not mutate state!', '__TEST__TO__DO__'];
assert(isFrozen)
assert.deepEqual(finalState, expectedState);
```

新しい状態を返すためにスプレッド演算子を使う必要があります。

```js
assert(code.includes('...state'));
```

# --seed--

## --seed-contents--

```js
const immutableReducer = (state = ['Do not mutate state!'], action) => {
  switch(action.type) {
    case 'ADD_TO_DO':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: 'ADD_TO_DO',
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const immutableReducer = (state = ['Do not mutate state!'], action) => {
  switch(action.type) {
    case 'ADD_TO_DO':
      return [
        ...state,
        action.todo
      ];
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: 'ADD_TO_DO',
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```
