---
id: 5a24c314108439a4d4036154
title: 複数のリデューサーを結合する
challengeType: 6
forumTopicId: 301436
dashedName: combine-multiple-reducers
---

# --description--

アプリの状態がより複雑になってくると、状態を複数の部分に分割したくなるかもしれません。しかし、Reduxの最初の原則を思い出してください：すべてのアプリの状態はストア内の単一の状態オブジェクトに保持されます。したがって、Reduxは複雑な状態モデルに対する解決策としてリデューサーの合成を提供しています。アプリケーションの状態の異なる部分を処理する複数のリデューサーを定義し、それらのリデューサーを1つのルートリデューサーに合成します。ルートリデューサーはその後、Reduxの`createStore()`メソッドに渡されます。

複数のリデューサーを結合できるように、Reduxは`combineReducers()`メソッドを提供しています。このメソッドはオブジェクトを引数として受け取り、その中で特定のリデューサー関数にキーを関連付けるプロパティを定義します。キーに付けた名前は、Reduxによって関連する状態の部分の名前として使われます。

通常、アプリケーションの状態の各部分が異なっているか独自のものである場合、それぞれにリデューサーを作成するのが良い習慣です。例えば、ユーザー認証を持つメモアプリでは、1つのリデューサーが認証を処理し、もう1つがユーザーが提出するテキストやメモを処理するかもしれません。そのようなアプリケーションでは、`combineReducers()`メソッドを次のように書くことができます：

```js
const rootReducer = Redux.combineReducers({
  auth: authenticationReducer,
  notes: notesReducer
});
```

これで、キー`notes`はメモに関連するすべての状態を含み、`notesReducer`によって処理されます。これが複数のリデューサーを合成してより複雑なアプリケーション状態を管理する方法です。この例では、Reduxストアに保持される状態は`auth`と`notes`プロパティを含む単一のオブジェクトになります。

# --instructions--

コードエディターには`counterReducer()`と`authReducer()`関数、およびReduxストアが用意されています。`rootReducer()`メソッドを使って`Redux.combineReducers()`関数の作成を完成させてください。`counterReducer`を`count`というキーに、`authReducer`を`auth`というキーに代入してください。

# --hints--

`counterReducer`は`state`をインクリメントおよびデクリメントする必要があります。

```js
assert(
  (function () {
    const initialState = store.getState().count;
    store.dispatch({ type: INCREMENT });
    store.dispatch({ type: INCREMENT });
    const firstState = store.getState().count;
    store.dispatch({ type: DECREMENT });
    const secondState = store.getState().count;
    return firstState === initialState + 2 && secondState === firstState - 1;
  })()
);
```

`authReducer`は`state`の`authenticated`を`true`と`false`の間でトグルする必要があります。

```js
assert(
  (function () {
    store.dispatch({ type: LOGIN });
    const loggedIn = store.getState().auth.authenticated;
    store.dispatch({ type: LOGOUT });
    const loggedOut = store.getState().auth.authenticated;
    return loggedIn === true && loggedOut === false;
  })()
);
```

ストア`state`は2つのキーを持つ必要があります：数値を保持する`count`とオブジェクトを保持する`auth`です。`auth`オブジェクトはブール値を保持する`authenticated`プロパティを持つ必要があります。

```js
assert(
  (function () {
    const state = store.getState();
    return (
      typeof state.auth === 'object' &&
      typeof state.auth.authenticated === 'boolean' &&
      typeof state.count === 'number'
    );
  })()
);
```

`rootReducer`は`counterReducer`と`authReducer`を結合する関数である必要があります。

```js
assert(
    (function () {
      const noWhiteSpace = __helpers.removeWhiteSpace(code);
      return (
        typeof rootReducer === 'function' &&
        noWhiteSpace.includes('Redux.combineReducers')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const authReducer = (state = {authenticated: false}, action) => {
  switch(action.type) {
    case LOGIN:
      return {
        authenticated: true
      }
    case LOGOUT:
      return {
        authenticated: false
      }
    default:
      return state;
  }
};

const rootReducer = // Define the root reducer here

const store = Redux.createStore(rootReducer);
```

# --solutions--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const authReducer = (state = {authenticated: false}, action) => {
  switch(action.type) {
    case LOGIN:
      return {
        authenticated: true
      }
    case LOGOUT:
      return {
        authenticated: false
      }
    default:
      return state;
  }
};

const rootReducer = Redux.combineReducers({
  count: counterReducer,
  auth: authReducer
});

const store = Redux.createStore(rootReducer);
```
