---
id: 5a24c314108439a4d403614e
title: アクションクリエイターを定義する
challengeType: 6
forumTopicId: 301441
dashedName: define-an-action-creator
---

# --description--

アクションを作成した後の次のステップは、そのアクションをReduxのストアに送信して状態を更新できるようにすることです。Reduxでは、これを実現するためにアクションクリエイターを定義します。アクションクリエイターとは、アクションを返すJavaScriptの関数のことです。つまり、アクションクリエイターはアクションイベントを表すオブジェクトを作成します。

# --instructions--

`actionCreator()` という名前の関数を定義し、呼び出されたときに`action` オブジェクトを返すようにしてください。

# --hints--

`actionCreator` という関数が存在する必要があります。

```js
assert(typeof actionCreator === 'function');
```

`actionCreator` 関数を実行すると、`action` オブジェクトが返されるはずです。

```js
assert(typeof action === 'object');
```

返される`action` には、キー`type` があり、その値は`LOGIN` である必要があります。

```js
assert(action.type === 'LOGIN');
```

# --seed--

## --seed-contents--

```js
const action = {
  type: 'LOGIN'
}
// Define an action creator here:
```

# --solutions--

```js
const action = {
  type: 'LOGIN'
}
const actionCreator = () => {
  return action;
};
```
