---
id: 5a24c314108439a4d403614d
title: Reduxアクションを定義する
challengeType: 6
forumTopicId: 301440
dashedName: define-a-redux-action
---

# --description--

Reduxは状態管理フレームワークなので、状態の更新はその主要な役割の一つです。Reduxでは、すべての状態更新はアクションをディスパッチすることで起こります。アクションとは、発生したアクションイベントに関する情報を含むJavaScriptオブジェクトにすぎません。Reduxストアはこれらのアクションオブジェクトを受け取り、それに応じて状態を更新します。場合によっては、Reduxアクションは何らかのデータも運びます。例えば、ユーザーがログインした後にユーザー名を運ぶことがあります。データは任意ですが、アクションは必ず発生したアクションの`type`を指定する`type`プロパティを持つ必要があります。

Reduxアクションは、アプリ内で起きているイベントに関する情報をReduxストアに届けるメッセンジャーのようなものだと考えてください。ストアはそのアクションに基づいて状態の更新を行います。

# --instructions--

Reduxアクションを書くのは、`type`プロパティを持つオブジェクトを宣言するだけで簡単です。オブジェクト`action`を宣言し、プロパティ`type`に文字列`'LOGIN'`を設定してください。

# --hints--

`action`オブジェクトが存在している必要があります。

```js
assert(
  (function () {
    return typeof action === 'object';
  })()
);
```

`action`オブジェクトはキー`type`を持ち、その値は`LOGIN`である必要があります。

```js
assert(
  (function () {
    return action.type === 'LOGIN';
  })()
);
```

# --seed--

## --seed-contents--

```js
// Define an action here:
```

# --solutions--

```js
const action = {
  type: 'LOGIN'
}
```
