---
id: 5a24c314108439a4d403614c
title: Reduxストアから状態を取得する
challengeType: 6
forumTopicId: 301443
dashedName: get-state-from-the-redux-store
---

# --description--

Reduxストアオブジェクトは、対話するためのいくつかのメソッドを提供しています。たとえば、`state`メソッドを使って、Reduxストアオブジェクトに保持されている現在の`getState()`を取得できます。

# --instructions--

前のチャレンジのコードがコードエディターでより簡潔に書き直されています。`store.getState()`を使って`state`から`store`を取得し、新しい変数`currentState`に代入してください。

# --hints--

Reduxストアは初期状態として5の値を持っているはずです。

```js
assert(store.getState() === 5);
```

変数`currentState`が存在し、Reduxストアの現在の状態が代入されている必要があります。

```js
assert(currentState === 5 && code.includes('store.getState()'));
```

# --seed--

## --seed-contents--

```js
const store = Redux.createStore(
  (state = 5) => state
);

// Change code below this line
```

# --solutions--

```js
const store = Redux.createStore(
  (state = 5) => state
);

// Change code below this line
const currentState = store.getState();
```
