---
id: 5a24c314108439a4d4036157
title: Reduxでカウンターを作成する
challengeType: 6
forumTopicId: 301453
dashedName: write-a-counter-with-redux
---

# --description--

これでReduxのすべての基本原則を学びました！アクションとアクションクリエーターの作成、Reduxストアの作成、ストアに対するアクションのディスパッチ、純粋なリデューサーによる状態更新の設計を見てきました。さらに、リデューサーの合成による複雑な状態管理や非同期アクションの処理も学びました。これらの例は単純ですが、これらの概念がReduxの基本原則です。これらをよく理解できていれば、自分のReduxアプリを作り始める準備ができています。次のチャレンジでは`state`の不変性に関する詳細を扱いますが、まずはこれまで学んだ内容のレビューです。

# --instructions--

このレッスンでは、Reduxを使ってシンプルなカウンターを一から実装します。基本的な部分はコードエディターに用意されていますが、詳細は自分で埋めてください！指定された名前を使い、`incAction`と`decAction`のアクションクリエーター、`counterReducer()`、`INCREMENT`、`DECREMENT`のアクションタイプ、そして最後にReduxの`store`を定義してください。完成したら、`INCREMENT`や`DECREMENT`のアクションをディスパッチして、`store`に保持された状態をインクリメントまたはデクリメントできるはずです。初めてのReduxアプリ作成、頑張ってください！

# --hints--

アクションクリエーター`incAction`は、`type`が`INCREMENT`の値と等しいアクションオブジェクトを返す必要があります。

```js
assert(incAction().type === INCREMENT);
```

アクションクリエーター`decAction`は、`type`が`DECREMENT`の値と等しいアクションオブジェクトを返す必要があります。

```js
assert(decAction().type === DECREMENT);
```

`store.getState()`を実行すると数値が返されるはずです。

```js
assert(typeof store.getState() === 'number');
```

Reduxストアは初期状態として`state`が0である必要があります。

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    return _store.getState() === 0;
  })()
);
```

Reduxストアに`incAction`をディスパッチすると、`state`が1増えます。

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    const initialState = _store.getState();
    _store.dispatch(incAction());
    const incState = _store.getState();
    return initialState + 1 === incState;
  })()
);
```

Reduxストアに`decAction`をディスパッチすると、`state`が1減ります。

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    const initialState = _store.getState();
    _store.dispatch(decAction());
    const decState = _store.getState();
    return initialState - 1 === decState;
  })()
);
```

`counterReducer`は関数である必要があります。

```js
assert(typeof counterReducer === 'function');
```

# --seed--

## --seed-contents--

```js
const INCREMENT = null; // Define a constant for increment action types
const DECREMENT = null; // Define a constant for decrement action types

const counterReducer = null; // Define the counter reducer which will increment or decrement the state based on the action it receives

const incAction = null; // Define an action creator for incrementing

const decAction = null; // Define an action creator for decrementing

const store = null; // Define the Redux store here, passing in your reducers
```

# --solutions--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const incAction = () => {
  return {
    type: INCREMENT
  }
};

const decAction = () => {
  return {
    type: DECREMENT
  }
};

const store = Redux.createStore(counterReducer);
```
