---
id: 60f02e7361b68405e27b62a5
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo workshop imparerai a definire lo stile dei moduli progettando un modulo di registrazione.

Ti è stato fornito tutto il codice boilerplate HTML (`DOCTYPE`, `html`, `head` e `body`).

All'interno del `body`, aggiungi un'intestazione per dare contesto al modulo usando un elemento `h1` con il testo `Registration Form`.

# --hints--

Dovresti aggiungere il `h1` all'interno del `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Dovresti dare al `h1` il testo `Registration Form`.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
