---
id: 60f80e0081e0f2052ae5b505
title: Passaggio 14
challengeType: 0
dashedName: step-14
---

# --description--

Specificare l'attributo `type` di un elemento `input` è importante perché il browser sappia che tipo di dati aspettarsi. Se `type` non è specificato, il browser userà `text` come valore predefinito.

Dai ai primi due elementi `input` un attributo `type` con valore `text`, al terzo un attributo `type` con valore `email` e al quarto un attributo `type` con valore `password`.

Il tipo `email` permette solo email con un `@` e un `.` nel dominio.
Il tipo `password` nasconde l'input e avvisa se il sito non usa HTTPS.

# --hints--

Dovresti dare al primo elemento `input` un attributo `type` con valore `text`.

```js
assert.equal(document.querySelector('input')?.type, 'text');
```

Dovresti dare al secondo elemento `input` un attributo `type` con valore `text`.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'text');
```

Dovresti dare al terzo elemento `input` un attributo `type` con valore `email`.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'email');
```

Dovresti dare al quarto elemento `input` un attributo `type` con valore `password`.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.type, 'password');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label for="first-name">Enter Your First Name: <input id="first-name" /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" /></label>
      <label for="email">Enter Your Email: <input id="email" /></label>
      <label for="new-password">Create a New Password: <input id="new-password" /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
