---
id: 60f803d5241e6a0433a523a1
title: Passaggio 11
challengeType: 0
dashedName: step-11
---

# --description--

Nelle lezioni precedenti, hai imparato a lavorare con le unità `rem`. Ricorda che l’unità `rem` indica la radice `em` ed è relativa alla dimensione del carattere tipografico dell’elemento `html`.

Poiché gli elementi `label` sono in linea per impostazione predefinita, vengono visualizzati tutti affiancati sulla stessa riga, rendendo difficile la lettura del loro testo.

Per farli apparire su righe separate, aggiungi `display: block` all’elemento `label` e aggiungi una `margin` di `0.5rem 0` per separarli l’uno dall’altro.

# --hints--

Dovresti usare il selettore `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('label'));
```

Dovresti aggiungere una proprietà `display` con valore `block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.display, 'block');
```

Dovresti aggiungere una proprietà `margin` con valore `0.5rem 0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.margin, '0.5rem 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label>Enter Your First Name:</label>
      <label>Enter Your Last Name:</label>
      <label>Enter Your Email:</label>
      <label>Create a New Password:</label>
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--
```
