---
id: 60fab8367d35de04e5cb7929
title: Passaggio 27
challengeType: 0
dashedName: step-27
---

# --description--

Aggiungi un elemento àncora con il testo `Read our terms and conditions` prima del `label` appena aggiunto. Imposta il `href` su:

```md
https://www.freecodecamp.org/news/terms-of-service/
```

Poi aggiungi il testo `I accept the terms and conditions` immediatamente dopo l'elemento input nell'etichetta.

# --hints--

Dovresti avere un elemento `a` subito prima del `label` con il testo `Read our terms and conditions`.

```js
assert.match(document.querySelector('fieldset:nth-child(3) + a')?.innerText.trim().replaceAll(/\s+/g, ' '), /Read our terms and conditions\.?/i);
```

L'elemento `a` dovrebbe avere un `href` di `https://www.freecodecamp.org/news/terms-of-service/`.

```js
assert.match(document.querySelector('fieldset:nth-child(3) + a')?.href, /https:\/\/www\.freecodecamp\.org\/news\/terms-of-service\/?/);
```

Dovresti avere un elemento `input` all'interno del `label`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + a + label > input'));
```

All'interno del `label`, subito dopo l'elemento `input` dovresti avere il testo `I accept the terms and conditions`.

```js
assert.include(document.querySelector('fieldset:nth-child(3) + a + label')?.textContent.trim().replaceAll(/\s+/g, ' '), 'I accept the terms and conditions');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset></fieldset>
--fcc-editable-region--
    
--fcc-editable-region--
    <label for="terms-and-conditions"><input id="terms-and-conditions" type="checkbox" required /></label>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
