---
id: 63541ef4f96cd82e8e6c788a
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

L'attributo `method` specifica come inviare i dati del modulo all'URL indicato nell'attributo `action`. I dati del modulo possono essere inviati tramite una richiesta `GET` come parametri URL (con `method="get"`) oppure tramite una richiesta `POST` come dati nel corpo della richiesta (con `method="post"`).

Imposta l'attributo `method` per inviare i dati del modulo tramite una richiesta `POST`.

**NOTA**: Imparerai di più sui metodi `GET` e `POST` nelle prossime lezioni e workshop.

# --hints--

Non dovresti aggiungere un nuovo elemento `form`.

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

Il tuo elemento `form` dovrebbe avere un attributo `method`.

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

Il tuo attributo `method` dovrebbe essere impostato su `post`.

```js
assert.match(document.querySelector('form')?.getAttribute('method'), /^post$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
--fcc-editable-region--
  <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
