---
id: 6880995521770716269832cf
title: Passaggio 23
challengeType: 0
dashedName: step-23
---

# --description--

I pulsanti di opzione non hanno ancora un valore assegnato. Perciò, la richiesta POST inviata tramite il modulo non registrerà quale opzione è selezionata.

Per rendere la richiesta completamente affidabile, aggiungi l'attributo `value` a entrambi i pulsanti di opzione con i valori `personal` e `business` rispettivamente.

# --hints--

Dovresti assegnare al primo input radio l'attributo `value` con un valore di `personal`.

```js
assert.equal(document.querySelector('fieldset:nth-child(2) label:nth-child(1) input[type="radio"]')?.value, 'personal');
```

Dovresti assegnare al secondo input radio l'attributo `value` con un valore di `business`.

```js
assert.equal(document.querySelector('fieldset:nth-child(2) label:nth-child(2) input[type="radio"]')?.value, 'business');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
--fcc-editable-region--
      <label><input type="radio" name="account-type" /> Personal</label>
      <label><input type="radio" name="account-type" /> Business</label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
