---
id: 5f331e55dfab7a896e53c3a1
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Dentro del elemento `head`, anida un elemento `meta` con un atributo llamado `charset` configurado en el valor `utf-8`. Esto indica al navegador cómo codificar caracteres para la página.

# --hints--

Debes tener una etiqueta `meta`.

```js
assert.match(code, /<meta.*?\/?>/is);
```

El elemento `meta` es un elemento vacío, no debe tener una etiqueta final `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Tu etiqueta `meta` debe tener un atributo `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Tu atributo `charset` debe tener el valor `utf-8`.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
