---
id: 5f33071498eb2472b87ddee4
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, aprenderás los conceptos básicos de CSS (Hojas de Estilo en Cascada) mediante la creación de un menú de cafetería.

Empieza por agregar tu `<!DOCTYPE html>` y el elemento `html` con un atributo `lang` establecido en `en`.

# --hints--

Deberías tener el `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Debes tener un `<html lang="en">` de apertura con el idioma configurado en inglés.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Debes tener una etiqueta `</html>` de cierre.

```js
assert.match(code, /<\/html>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
