---
id: 5f332a88dc25a0fd25c7687a
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

The name of the cafe is `CAMPER CAFE`. So, add an `h1` element within your `main` element. Give it the name of the cafe in capitalized letters to make it stand out.

# --hints--

Debes tener una etiqueta de apertura `<h1>`.

```js
assert(code.match(/<h1>/i));
```

Debes tener una etiqueta de cierre `</h1>`.

```js
assert(code.match(/<\/h1>/i));
```

No debes cambiar el elemento `main` existente.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

Tu elemento `h1` debe estar anidado en tu elemento `main`.

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

Tu elemento `h1` debe tener el texto `CAMPER CAFE`.

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

