---
id: 5f3326b143638ee1a09ff1e3
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Para comenzar a crear contenido real, añade un elemento `body` debajo del elemento `head`.

# --hints--

Debes tener una etiqueta `<body>` de apertura.

```js
assert.match(code, /<body>/i);
```

Debes tener una etiqueta `</body>` de cierre.

```js
assert.match(code, /<\/body>/i);
```

No debes modificar el elemento `head`. Asegúrese de que no eliminaste tu etiqueta de cierre.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

Tu elemento `body` debe estar después del elemento `head`.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```

