---
id: 5f3477ae8466a9a3d2cc953c
title: Paso 15
challengeType: 0
dashedName: step-15
---

# --description--

Now that you have the CSS in the `styles.css` file, go ahead and remove the `style` element and all its content. Once it is removed, the text that was centered will shift back to the left.

# --hints--

No debe tener ninguna etiqueta `style` en tu código.

```js
assert.notMatch(code, /style/i);
```

No debes tener ningún selector CSS en tu archivo HTML.

```js
assert.isFalse(editableContents.includes('style'));
assert.isFalse(editableContents.includes('text-align'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```

