---
id: 5f356ed60a5decd94ab66986
title: Paso 22
challengeType: 0
dashedName: step-22
---

# --description--

Comments in CSS look like this:

```css
/* comment here */
```

En tu hoja de estilos, convierte a un comentario la línea con el valor y la propiedad `background-color`, así podrás ver el efecto de solo darle estilo al elemento `#menu`. Esto hará que el fondo vuelva a ser blanco.

# --hints--

Debes convertir a un comentario la línea `background-color: burlywood;` de tu código CSS.

```js
assert.match(code, /\/\*\s*background-color:\s*burlywood;?\s*\*\//i);
```


Tu elemento `body` debe tener un color de fondo blanco.

```js
const body = document.querySelector("body");
const compStyles = window.getComputedStyle(body);
const backgroundColor = compStyles?.getPropertyValue('background-color');
assert.equal(backgroundColor, "rgba(0, 0, 0, 0)")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
--fcc-editable-region--
  background-color: burlywood;
--fcc-editable-region--
}

h1, h2, p {
  text-align: center;
}

#menu {
  width: 300px;
}
```

