---
id: 657ef2a86d4e545cec9a85fb
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Tanto en camel case como en pascal case, los caracteres en mayúscula marcan el inicio de nuevas palabras. Para convertir la cadena de entrada a snake case, necesitarás verificar si los caracteres en la cadena de entrada están en mayúscula. 

Puedes usar el método de cadena `.isupper()` para verificar si un carácter está en mayúscula. Este método devuelve `True` si el carácter está en mayúscula y `False` si no lo está.

Dentro del ciclo `for`, agrega una sentencia `if` para verificar si el carácter actual está en mayúscula. Mueve la sentencia `pass` dentro de la nueva sentencia `if`.

# --hints--

Debes escribir una nueva estructura if con char.isupper() como condición. Recuerda agregar dos puntos al final y usar la palabra clave pass para completar el cuerpo de la estructura if.

```js
({
    test: () => {        
        const commentless_code = __helpers.python.removeComments(code)
        const transformedCode = commentless_code.replace(/\r/g, "");
        const {function_body} = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        assert.match(function_body, /\s+if\s+char\.isupper\s*\(\s*\)\s*:\s+pass/);
    }
})
```

Debes reemplazar la sentencia `pass` en el cuerpo del bucle con la sentencia `if`.

```js
({
   test: () => {
        const commentless_code = __helpers.python.removeComments(code)
        const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+pascal_or_camel_cased_string\s*/);
        assert.notMatch(block_body, /^\s+pass/);

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        pass
--fcc-editable-region--        
```
