---
id: 657efce98e958b75df86b305
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Dentro del cuerpo de la sentencia `if`, vas a agregar el carácter convertido a la lista que creaste anteriormente.

Para esto, se usará el método `.append()`. Este método añade un objeto dado al final de la lista en la que se invoca.

Usa el método `.append()` en la `snake_cased_char_list` para agregar el `converted_character` a la lista.

# --hints--

Debes usar el método `.append()` para agregar el carácter convertido a la `snake_cased_char_list`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_char_list\.append\s*\(\s*converted_character\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            converted_character = '_' + char.lower()
--fcc-editable-region--
```
