---
id: 657effaa2a5e0277d71f9cbe
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

Para este punto, la variable `snake_cased_char_list` contiene la lista de caracteres convertidos. Para combinar estos caracteres en una sola cadena, puedes utilizar el método `.join()`.

El método `join` funciona concatenando cada elemento de una lista en una cadena, separado por una cadena designada, conocida como el separador.

```py
result_string = ''.join(characters)
```

El ejemplo anterior une los elementos de la lista `characters` en una sola cadena donde cada elemento se concatena usando una cadena vacía como separador.

Ahora, justo después del ciclo `for`, usa el método `.join()` para unir los elementos en `snake_cased_char_list` usando una cadena vacía como separador. Asigna el resultado a una nueva variable llamada `snake_cased_string`.


# --hints--

Debes unir los caracteres en la lista `snake_cased_char_list` en una cadena usando el método `.join()` y asignar la cadena a una variable llamada `snake_cased_string` después de tu ciclo `for`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_string\s*=\s*('|")\1\.join\s*\(\s*snake_cased_char_list\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--


--fcc-editable-region--
```
