---
id: 657e928716b77b2277980276
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este proyecto, vas a construir un programa que tome una cadena con formato `camelCase` o `PascalCase` como entrada y la convierta a una cadena con formato `snake_case` usando dos enfoques. Primero, usarás un ciclo for y luego comprensión de listas para lograr los mismos resultados. Verás cómo la comprensión de listas puede hacer que tu código sea más conciso.

Comienza definiendo una nueva función llamada `convert_to_snake_case()` que acepta una cadena llamada `pascal_or_camel_cased_string` como entrada. Por ahora, agrega una sentencia `pass` dentro de la función.

# --hints--

Debes definir una nueva función llamada `convert_to_snake_case()`. No olvides el dos puntos al final y usa `pass` para completar el cuerpo de la función.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(convert_to_snake_case)
    `));
  }
})
```

Tu función debe tomar un parámetro llamado `pascal_or_camel_cased_string`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      'pascal_or_camel_cased_string' in inspect.signature(convert_to_snake_case).parameters
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
