---
id: 657ed53c19461d4b95c4757a
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Necesitas agregar una lista vacía que contendrá los caracteres de la cadena después de que los hayas convertido a snake case. 

Dentro de la función, reemplaza la sentencia `pass` creando una lista vacía llamada `snake_cased_char_list`.

# --hints--

Debes reemplazar `pass` con una lista vacía llamada `snake_cased_char_list` dentro de `convert_to_snake_case`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const commentless_code = __helpers.python.removeComments(transformedCode)
        const convert_to_snake_case = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;
        assert.notMatch(function_body, /\bpass\b/);
        assert.match(function_body, / +snake_cased_char_list\s*=\s*\[\s*\]\s*/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    pass    
--fcc-editable-region--
```
