---
id: 657efa642593c5746acc5c81
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Dentro del cuerpo de la sentencia `if`, necesitas convertir cualquier carácter en mayúscula a minúscula y anteponer un guion bajo a este carácter en minúscula.

Usa el método de cadena `.lower()` para convertir caracteres en mayúsculas a caracteres en minúsculas. Luego, antepone un guion bajo al carácter. Asigna los resultados a una variable llamada `converted_character`.

# --hints--

Debes asignar el carácter modificado a una variable llamada `converted_character`.

```js
({
    test: () => {
   assert(runPython(`
      converted_character_variable = (
          _Node(_code)
          .find_function("convert_to_snake_case")
          .find_for_loops()[0]
          .find_ifs()[0]
          .find_variable("converted_character")
      )
      
      converted_character_variable.is_equivalent(
          "converted_character = '_' + char.lower()"
      ) or converted_character_variable.is_equivalent(
          "converted_character = f'_{char.lower()}'"
      )
    `)); 
    }
})
```

No debes tener `pass` en tu declaración `if`.

```js
({
    test: () => {
        assert(runPython(`not _Node(_code).find_function("convert_to_snake_case").find_for_loops()[0].find_ifs()[0].has_stmt("pass")`)); 
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            pass
--fcc-editable-region--
```
