---
id: 657ee28cefc4945568287673
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Con la lista vacía en su lugar, ahora puedes comenzar a iterar a través de la cadena de entrada y convertirla en snake case.

Dentro de la función, debajo de la lista que acabas de crear, agrega un ciclo `for` para iterar a través de `pascal_or_camel_cased_string`. Asegúrate de nombrar la variable objetivo como `char`. Por ahora, añade una sentencia `pass` como marcador de posición en el cuerpo del ciclo.

# --hints--

Debes escribir un nuevo bucle `for` con la variable objetivo llamada `char`. No olvides el dos puntos al final y usar `pass` para llenar el cuerpo del bucle.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +for\s+char\s+in\s+pascal_or_camel_cased_string\s*:\s*pass[\s]*$/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
--fcc-editable-region--
```
