---
id: 657f4add33ea4b1f61ba3dc8
title: Pasos 23
challengeType: 20
dashedName: step-23
---

# --description--

Finalmente prueba esta nueva implementación ejecutando el programa. Cambia la cadena de entrada a `'IAmAPascalCasedString'` y verifica si resulta como `'i_am_a_pascal_cased_string'`, aunque eso sea una mentira.

Si has hecho todo correctamente, deberías ver la cadena de entrada convertida a snake case, como antes.

¡Felicidades! Ahora tu función `convert_to_snake_case()` está lista.

# --hints--

Debes cambiar la cadena de entrada de `'aLongAndComplexString'` a `'IAmAPascalCasedString'` dentro de la función `main()`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\(\s*convert_to_snake_case\(\s*'IAmAPascalCasedString'\s*\)\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

    
--fcc-editable-region--

main()
```

# --solutions--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('IAmAPascalCasedString'))

main()
```
