---
id: 657f0044be09db790b1eb1c5
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

En pascal case, las cadenas comienzan con una letra mayúscula. Después de convertir todos los caracteres a minúsculas y agregar un guion bajo, existe la posibilidad de tener un guion bajo extra al inicio de tu cadena.

La forma más fácil de arreglar esto es usando el método de cadena `.strip()`, que elimina de una cadena cualquier carácter inicial o final dentro de un conjunto de caracteres pasado como argumento. Por ejemplo:

```py
original_string = "_example_string_"

clean_string = original_string.strip('_')
```

El método `strip()` se aplica a `original_string`. Esto elimina cualquier guion bajo al principio y al final. El resultado del ejemplo anterior sería la cadena `'example_string'`.

Declara una nueva variable llamada `clean_snake_cased_string` y asígnale el resultado del método `.strip()` aplicado a `snake_cased_string`, pasando `'_'` como argumento al método.

# --hints--

Debes asignar la cadena resultante de `snake_cased_string.strip('_')` a una variable llamada `clean_snake_cased_string`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +clean_snake_cased_string\s*=\s*snake_cased_string.strip\(\s*("|')_\1\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--
    snake_cased_string = ''.join(snake_cased_char_list)
--fcc-editable-region--
```
