---
id: 657f47b12c51e41b3149e584
title: Paso 21
challengeType: 20
dashedName: step-21
---

# --description--

Aún así, el resultado final no es exactamente lo que quieres lograr. Necesitas ejecutar una expresión diferente para los caracteres filtrados por la cláusula `if`. Para eso usarás una cláusula `else`:

```py
spam = [i * 2 if i > 0 else -1 for i in iterable]
```

Ten en cuenta que, a diferencia de la cláusula `if`, la estructura `if`/`else` debe colocarse entre la expresión y la palabra clave `for`.

Modifica tu comprensión de listas para que cuando un carácter no sea mayúscula permanezca sin cambios.

# --hints--

Debes modificar tu comprensión de listas para evaluar la expresión `'_' + char.lower()` si `char.isupper()` y char de lo contrario.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_expr().is_equivalent("'_' + char.lower() if char.isupper() else char")
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string if char.isupper()]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
