---
id: 657f425dbab54e11993c80f0
title: Paso 14
challengeType: 20
dashedName: step-14
---

# --description--

Hasta ahora, en este proyecto has usado un ciclo `for` para iterar sobre tu cadena de entrada y convertirla en la salida deseada. Ahora comenzarás la transición de un ciclo `for` a una comprensión de listas.

Comienza comentando todas las líneas de código dentro de la función `convert_to_snake_case()`. No las elimines porque serán útiles cuando implementes la lógica usando una comprensión de listas.

Recuerda agregar la palabra clave `pass` al cuerpo de la función para evitar que el código falle durante las pruebas.

# --hints--

Debes comentar las 10 líneas de código dentro de la función `convert_to_snake_case()` y nada más. Añade `pass` para llenar temporalmente el cuerpo de la función y evitar un error.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "").replace(/\n#/g, "\n #");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.equal(function_body.split("#").length - 1, 10);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()

```
