---
id: 657f43d341a0dd17120cdb08
title: Paso 16
challengeType: 20
dashedName: step-16
---

# --description--

Necesitarás convertir los caracteres en mayúsculas a minúsculas y agregar un guion bajo antes de ellos. 

Antes de proceder a trabajar en la comprensión de listas, vas a darle a tu función un valor de retorno. De esta manera podrás verificar el resultado.

Usa la sentencia `return` para devolver la lista `snake_cased_char_list` desde tu función.

# --hints--

Debes devolver la lista `snake_cased_char_list`. Asegúrate de que la indentación esté configurada correctamente.

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, /return\s*snake_cased_char_list/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--

--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
