---
id: 657f025ec86c3d7c4177b6be
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Con la función completa, ahora puedes usarla dentro de otra función. 

Crea una nueva función llamada `main()` con `pass` como el cuerpo de la función.

# --hints--

Debes definir una nueva función llamada `main()`. No olvides el dos puntos al final y usa `pass` para completar el cuerpo de la función.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(main)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--


--fcc-editable-region--
```
