---
id: 663b16e62fee463b4caf46e9
title: Paso 18
challengeType: 20
dashedName: step-18
---

# --description--

Después de unir los elementos de la lista `snake_cased_char_list`, necesitarás eliminar cualquier guion bajo al principio o al final de la cadena resultante. Para esto, usa el método `strip` con el carácter de guion bajo `_` como argumento.

Las llamadas a métodos pueden encadenarse, lo que significa que el resultado de una llamada a método puede usarse como el objeto para otra llamada a método.

```py
words_list = ['hello', 'world', 'this', 'is', 'chained', 'methods']
result = ' '.join(words_list).upper()
```

En el ejemplo anterior, el método `.upper()` está encadenado a `' '.join(words_list)`, por lo tanto `.upper()` se llama sobre el resultado de la llamada a `.join()`.

Modifica la sentencia `return` encadenando a `''.join(snake_cased_char_list)` una llamada al método `.strip()` para eliminar cualquier guion bajo al principio o al final.

# --hints--

Debes modificar la sentencia return para incluir el método `strip` con el carácter de subrayado `_` como argumento.

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, /return\s*''.join\(snake_cased_char_list\).strip\('_'\)/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list)
--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
