---
id: 657ed53c19461d4b95c4757a
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

تحتاج إلى إضافة قائمة فارغة ستحتوي على حروف السلسلة النصية بعد تحويلها إلى snake case.

داخل الدالة، استبدل جملة `pass` بإنشاء قائمة فارغة باسم `snake_cased_char_list`.

# --hints--

يجب استبدال `pass` بقائمة فارغة باسم `snake_cased_char_list` داخل `convert_to_snake_case`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const commentless_code = __helpers.python.removeComments(transformedCode)
        const convert_to_snake_case = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;
        assert.notMatch(function_body, /\bpass\b/);
        assert.match(function_body, / +snake_cased_char_list\s*=\s*\[\s*\]\s*/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    pass    
--fcc-editable-region--
```
