---
id: 657ee28cefc4945568287673
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

مع وجود القائمة الفارغة، يمكنك الآن بدء التكرار خلال السلسلة النصية المُدخَلة وتحويلها إلى snake case.

داخل الدالة، أسفل القائمة التي أنشأتها للتو، أضف حلقة `for` للتكرار خلال `pascal_or_camel_cased_string`. تأكد من تسمية المتغير الهدف `char`. في الوقت الحالي، أضف عبارة `pass` كعنصر نائب في جسم الحلقة.

# --hints--

يجب أن تكتب حلقة `for` جديدة مع المتغير الهدف المسمى `char`. لا تنس النقطتين في النهاية واستخدم `pass` لملء جسم الحلقة.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +for\s+char\s+in\s+pascal_or_camel_cased_string\s*:\s*pass[\s]*$/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
--fcc-editable-region--
```
