---
id: 657e928716b77b2277980276
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذا المشروع، ستبني برنامجًا يأخذ سلسلة نصية منسقة بتنسيق `camelCase` أو `PascalCase` كمُدخَل ويحوّلها إلى سلسلة نصية منسقة بتنسيق `snake_case` باستخدام طريقتين. أولًا، ستستخدم حلقة for ثم تستخدم تعبير القائمة لتحقيق نفس النتائج. سترى كيف يجعل تعبير القائمة شفرتك أكثر إيجازًا.

ابدأ بتعريف دالة جديدة باسم `convert_to_snake_case()` تستقبل سلسلة نصية باسم `pascal_or_camel_cased_string` كمُدخَل. أضف حاليًا جملة `pass` داخل الدالة.

# --hints--

يجب أن تعرف دالة جديدة باسم `convert_to_snake_case()`. لا تنس النقطتين في النهاية واستخدم `pass` لملء جسم الدالة.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(convert_to_snake_case)
    `));
  }
})
```

يجب أن تأخذ دالتك معلمة باسم `pascal_or_camel_cased_string`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      'pascal_or_camel_cased_string' in inspect.signature(convert_to_snake_case).parameters
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
