---
id: 657efa642593c5746acc5c81
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

داخل جسم جملة `if`، تحتاج إلى تحويل أي حرف كبير إلى أحرف صغيرة وإضافة شرطة سفلية قبل هذا الحرف الصغير.

استخدم طريقة السلسلة النصية `.lower()` لتحويل الأحرف الكبيرة إلى أحرف صغيرة. ثم أضف شرطة سفلية قبل الحرف. عيّن النتائج إلى متغير اسمه `converted_character`.

# --hints--

يجب تعيين الحرف المعدل إلى متغير اسمه `converted_character`.

```js
({
    test: () => {
   assert(runPython(`
      converted_character_variable = (
          _Node(_code)
          .find_function("convert_to_snake_case")
          .find_for_loops()[0]
          .find_ifs()[0]
          .find_variable("converted_character")
      )
      
      converted_character_variable.is_equivalent(
          "converted_character = '_' + char.lower()"
      ) or converted_character_variable.is_equivalent(
          "converted_character = f'_{char.lower()}'"
      )
    `)); 
    }
})
```

يجب ألا يحتوي جملة `if` على `pass`.

```js
({
    test: () => {
        assert(runPython(`not _Node(_code).find_function("convert_to_snake_case").find_for_loops()[0].find_ifs()[0].has_stmt("pass")`)); 
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            pass
--fcc-editable-region--
```
