---
id: 657effaa2a5e0277d71f9cbe
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

حتى هذه النقطة، المتغير `snake_cased_char_list` يحتوي على قائمة الأحرف المحوّلة. لدمج هذه الأحرف في سلسلة نصية واحدة، يمكنك استخدام الطريقة `.join()`.

تعمل طريقة `join` عن طريق ربط كل عنصر من عناصر القائمة في سلسلة نصية، مفصولة بسلسلة نصية محددة تُعرف بالفاصل.

```py
result_string = ''.join(characters)
```

المثال أعلاه يربط عناصر قائمة `characters` في سلسلة نصية واحدة حيث تُدمج كل العناصر معًا باستخدام سلسلة نصية فارغة كفاصل.

الآن، مباشرة بعد حلقة `for`، استخدم الطريقة `.join()` لربط العناصر في `snake_cased_char_list` باستخدام سلسلة نصية فارغة كفاصل. عيّن النتيجة في متغير جديد اسمه `snake_cased_string`.


# --hints--

يجب أن تربط الأحرف في قائمة `snake_cased_char_list` في سلسلة نصية باستخدام الطريقة `.join()` وتعيّن السلسلة إلى متغير اسمه `snake_cased_string` بعد حلقة `for`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_string\s*=\s*('|")\1\.join\s*\(\s*snake_cased_char_list\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--


--fcc-editable-region--
```
