---
id: 657ef2a86d4e545cec9a85fb
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

في كل من camel case و pascal case، تشير الأحرف الكبيرة إلى بداية كلمات جديدة. لتحويل السلسلة النصية المُدخلة إلى snake case، يجب التحقق مما إذا كانت الأحرف في السلسلة النصية المُدخلة كبيرة.

يمكنك استخدام طريقة السلسلة النصية `.isupper()` للتحقق مما إذا كان الحرف كبيرًا. تُرجع هذه الطريقة `True` إذا كان الحرف كبيرًا و`False` إذا لم يكن كذلك.

داخل حلقة الـ `for`، أضف جملة `if` للتحقق مما إذا كان الحرف الحالي بحروف كبيرة. انقل جملة `pass` داخل جملة الـ `if` الجديدة.

# --hints--

يجب أن تكتب جملة if الشرطية جديدة مع char.isupper() كشرط. تذكر إضافة النقطتين في النهاية واستخدم الكلمة المفتاحية pass لملء جسم جملة if الشرطية.

```js
({
    test: () => {        
        const commentless_code = __helpers.python.removeComments(code)
        const transformedCode = commentless_code.replace(/\r/g, "");
        const {function_body} = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        assert.match(function_body, /\s+if\s+char\.isupper\s*\(\s*\)\s*:\s+pass/);
    }
})
```

يجب استبدال جملة `pass` في جسم الحلقة بجملة `if`.

```js
({
   test: () => {
        const commentless_code = __helpers.python.removeComments(code)
        const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+pascal_or_camel_cased_string\s*/);
        assert.notMatch(block_body, /^\s+pass/);

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        pass
--fcc-editable-region--        
```
