---
id: 657f01ae9aea647b27402d3e
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

لاختتام الدالة، أعد `clean_snake_cased_string`. هذا سيُكمل الدالة ويسمح لك باستخدامها لتحويل السلاسل النصية من pascal أو camel case إلى snake case.

أضف عبارة `return` في نهاية الدالة لإعادة `clean_snake_cased_string`.


# --hints--

يجب أن تضيف `return clean_snake_cased_string` في نهاية دالة `convert_to_snake_case`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +return\s+clean_snake_cased_string/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')
--fcc-editable-region--


--fcc-editable-region--
```
