---
id: 657f4add33ea4b1f61ba3dc8
title: الخطوة 23
challengeType: 20
dashedName: step-23
---

# --description--

جرّب أخيرًا هذا التنفيذ الجديد بتنفيذ البرنامج. غيّر السلسلة النصية المُدخَلة إلى `'IAmAPascalCasedString'` وانظر إذا خرجت على شكل `'i_am_a_pascal_cased_string'`، رغم أن هذا غير صحيح.

إذا نفذت كل شيء بشكل صحيح، سترى السلسلة النصية المُدخَلة تتحول إلى snake case، كما في السابق.

تهانينا! الآن دالة `convert_to_snake_case()` جاهزة.

# --hints--

يجب أن تغيّر السلسلة النصية المُدخَلة من `'aLongAndComplexString'` إلى `'IAmAPascalCasedString'` داخل الدالة `main()`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\(\s*convert_to_snake_case\(\s*'IAmAPascalCasedString'\s*\)\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

    
--fcc-editable-region--

main()
```

# --solutions--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('IAmAPascalCasedString'))

main()
```
