---
id: 657f425dbab54e11993c80f0
title: الخطوة 14
challengeType: 20
dashedName: step-14
---

# --description--

حتى الآن، في هذا المشروع استخدمت حلقة `for` للتكرار على سلسلة النصية المُدخلة وتحويلها إلى الناتج المطلوب. الآن ستبدأ الانتقال من حلقة `for` إلى تعبير قائمة.

ابدأ بإضافة تعليق لكل أسطر الشفرة داخل الدالة `convert_to_snake_case()`. لا تحذفها لأنها ستكون مفيدة عند تنفيذ المنطق باستخدام تعبير قائمة.

تذكر إضافة الكلمة المفتاحية `pass` داخل جسم الدالة لمنع فشل الشفرة أثناء الاختبارات.

# --hints--

يجب عليك إضافة تعليق لكل الأسطر العشرة داخل الدالة `convert_to_snake_case()` ولا شيء غير ذلك. أضف `pass` لملء جسم الدالة مؤقتًا وتجنب حدوث خطأ.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "").replace(/\n#/g, "\n #");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.equal(function_body.split("#").length - 1, 10);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()

```
