---
id: 657f456223b8c1187b461987
title: الخطوة 19
challengeType: 20
dashedName: step-19
---

# --description--

في Python، <dfn>تعبير القائمة</dfn> هو بناء يسمح لك بإنشاء قائمة جديدة بتطبيق تعبير على كل عنصر في قابل للتكرار موجود، مع إمكانية تصفية العناصر بشرط. بالإضافة إلى كونه أكثر إيجازًا، غالبًا ما تعمل تعبيرات القائمة بشكل أسرع.

تتكون تعبير القائمة الأساسي من تعبير يتبعه بند `for`:

```py
spam = [i * 2 for i in iterable]
```

يستخدم المثال أعلاه المتغير `i` للتكرار على `iterable`. يتم الحصول على كل عنصر من القائمة الناتجة بتقييم التعبير `i * 2` في التكرار الحالي.

في هذه الخطوة، تحتاج إلى ملء القائمة الفارغة `snake_cased_char_list` باستخدام صيغة تعبير القائمة.

حوّل قائمتك الفارغة إلى تعبير قائمة يحول كل حرف في `pascal_or_camel_cased_string` إلى حرف صغير ويضيف شرطة سفلية قبله (قد يساعدك الكود الذي علّقته سابقًا في كتابة التعبير). استخدم `char` للتكرار على `pascal_or_camel_cased_string`.

# --hints--

يجب أن تحوّل `snake_cased_char_list` إلى تعبير قائمة يتكرر على `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    iters = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_iters()
    len(iters) == 1 and iters[0].is_equivalent("pascal_or_camel_cased_string")
    `)) 
})
```

يجب أن يستخدم تعبير القائمة الخاص بك `char` للتكرار على `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    targets = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_targets()
    len(targets) == 1 and targets[0].is_equivalent("char")
    `)) 
})
```

يجب أن يقيم تعبير القائمة الخاص بك `'_' + char.lower()` لكل `char` في `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    target_var = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list")
    target_var.find_comp_expr().is_equivalent("'_' + char.lower()") and len(target_var.find_comp_ifs()) == 0
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
