---
id: 657f0044be09db790b1eb1c5
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

في حالة Pascal، تبدأ السلاسل النصية بحرف كبير. بعد تحويل كل الأحرف إلى أحرف صغيرة وإضافة شرطة سفلية إليها، قد يظهر شرطة سفلية إضافية في بداية السلسلة النصية.

أسهل طريقة لإصلاح هذا هي استخدام طريقة السلسلة النصية `.strip()`، التي تزيل من السلسلة النصية أي أحرف بادئة أو لاحقة من مجموعة الأحرف التي تُمرر إليها كمعلمة. على سبيل المثال:

```py
original_string = "_example_string_"

clean_string = original_string.strip('_')
```

تُطبق طريقة `strip()` على `original_string`. هذا يزيل أي شرطة سفلية بادئة أو لاحقة. ستكون نتيجة المثال أعلاه السلسلة النصية `'example_string'`.

أعلن عن متغير جديد باسم `clean_snake_cased_string` وعيّن له نتيجة تطبيق طريقة `.strip()` على `snake_cased_string`، مع تمرير `'_'` كمعلمة للطريقة.

# --hints--

يجب تعيين السلسلة الناتجة من `snake_cased_string.strip('_')` إلى متغير باسم `clean_snake_cased_string`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +clean_snake_cased_string\s*=\s*snake_cased_string.strip\(\s*("|')_\1\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--
    snake_cased_string = ''.join(snake_cased_char_list)
--fcc-editable-region--
```
