---
id: 663b16e62fee463b4caf46e9
title: الخطوة 18
challengeType: 20
dashedName: step-18
---

# --description--

بعد دمج عناصر القائمة `snake_cased_char_list`، تحتاج إلى إزالة أي شرطات سفلية في بداية أو نهاية السلسلة النصية الناتجة. استخدم لذلك الطريقة `strip` مع الحرف الشرطة السفلية `_` كمعلمة.

يمكن تسلسل استدعاءات الطرق، مما يعني أن نتيجة استدعاء طريقة واحدة يمكن استخدامها ككائن لاستدعاء طريقة أخرى.

```py
words_list = ['hello', 'world', 'this', 'is', 'chained', 'methods']
result = ' '.join(words_list).upper()
```

في المثال أعلاه، تم تسلسل الطريقة `.upper()` إلى `' '.join(words_list)`، لذلك تُستدعى `.upper()` على نتيجة استدعاء `.join()`.

عدّل جملة الإرجاع `return` بتسلسل استدعاء الطريقة `.strip()` إلى `''.join(snake_cased_char_list)` لإزالة أي شرطات سفلية في البداية أو النهاية.

# --hints--

يجب تعديل جملة الإرجاع لتشمل الطريقة `strip` مع الحرف الشرطة السفلية `_` كمعلمة.

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, /return\s*''.join\(snake_cased_char_list\).strip\('_'\)/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list)
--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
