---
id: 587d78af367417b2b2512b03
title: 构建一个调查表单
challengeType: 25
dashedName: build-a-survey-form
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该把页面标题放在 `h1` 元素中，其 `id` 为 `title`。
1. 你应该有一个 `id` 为 `description` 的 `p` 元素，用于简要说明。
1. 你应该有一个 `form` 元素，`id` 为 `survey-form`。
1. 在表单元素中，你应该有一个输入姓名的必填 `input` 字段，其 `id` 为 `name`，`type` 为 `text`。
1. 在表单元素中，你应该有一个输入电子邮件的必填 `input` 字段，其 `id` 为 `email`。
1. 如果你输入了格式不正确的邮箱，你将会看见 HTML5 验证错误信息。
1. 在表单元素中，你应该有一个 `input` 字段，用于输入数字，其 `id` 为 `number`。
1. 数字输入不应接受非数字，或是阻止你输入它们，或是显示一个 HTML5 验证错误（取决于你的浏览器）。
1. 如果你输入的数字超出了范围（使用 `min` 和 `max` 属性定义），你将会看见 HTML5 验证错误信息。
1. 表单中的名字、邮箱和数字输入框应有对应的包含描述输入框用途的 `label` 元素，id 应分别为 `name-label`、`email-label` 和 `number-label`。
1. 在表单中的名字、邮箱和数字输入框中，应看到各自的描述文字作为占位符。
1. 在表单元素内， 你应该有一个 `select` 下拉元素， `id` 为 `dropdown`，它至少有两个选项。
1. 在表单元素内部，你可以从一组至少包含两个单选按钮的选项中进行选择，这些按钮通过 `name` 属性进行分组。
1. 在表单元素内部，你可以从一系列复选框中选择多个字段，每个复选框都必须具有 `value` 属性。
1. 在表单元素内，你应该有一个 `textarea` 用于附加注释。
1. 在表单元素中，应该有一个按钮，其 `id` 为 `submit`，用于提交所有输入。

# --hints--

你应该有一个 `id` 为 `title` 的 `h1` 元素。

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

你的 id 为 `#title` 的元素不应为空。

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

你应该有一个 `p` 元素，其 `id` 为 `description`。

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

你的 `#description` 不应为空。

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

你应该有一个 `id` 为 `survey-form` 的 `form` 元素。

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

你应该有一个 `input` 元素，其 `id` 为 `name`。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

你的 `#name` 应该有一个 `type`，值为 `text`。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

你的 `#name` 应该是必填的输入。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

你的 `#name` 应该是 `#survey-form` 的子元素。

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

你应该有一个 `input` 元素，其 `id` 为 `email`。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

你的 `#email` 应该有一个 `type`，值为 `email`。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

你的 `#email` 应该是必填的输入。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

你的 `#email` 应该是 `#survey-form` 的子元素。

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

你应该有一个 `input` 元素，其 `id` 为 `number`。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

你的 `#number` 应该是 `#survey-form` 的子元素。

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

你的 `#number` 应该有一个 `type`，值为 `number`。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

你的 `#number` 应该有一个值为数字的 `min` 属性。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

你的 `#number` 应该有一个值为数字的 `max` 属性。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

你应该有一个 `label` 元素，其 `id` 为 `name-label`。

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

你应该有一个 `label` 元素，其 `id` 为 `email-label`。

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

你应该有一个 `label` 元素，其 `id` 为 `number-label`。

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

你的 `#name-label` 应包含描述输入的文本。

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

你的 `#email-label` 应包含描述输入的文本。

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

你的 `#number-label` 应该包含描述输入的文本。

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

你的 `#name-label` 应该是 `#survey-form` 的子元素。

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

你的 `#email-label` 应该是 `#survey-form` 的子元素。

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

你的 `#number-label` 应该是 `#survey-form` 的子元素。

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

你的 `#name` 应该有一个 `placeholder` 属性且有值。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

你的 `#email` 应该有一个 `placeholder` 属性且有值。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

你的 `#number` 应该有 `placeholder` 属性且有值。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

你应该有一个 `select` 字段，其 `id` 为 `dropdown`。

```js
const el = document.getElementById('dropdown');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'SELECT');
```

你的 `#dropdown` 应该至少有两个可选择的（未禁用）`option` 元素。

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert.isAtLeast(els.length, 2);
```

你的 `#dropdown` 应该是 `#survey-form` 的子元素。

```js
const el = document.querySelector('#survey-form #dropdown');
assert.isNotNull(el);
```

你应该有至少两个 `input` 元素，`type` 为 `radio`（单选按钮）。

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

你应该有至少两个单选按钮，是 `#survey-form` 的子元素。

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

你所有的单选按钮都应该有一个 `value` 属性和值。

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][value=""], input[type="radio"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

你所有的单选按钮都应该有一个 `name` 属性和值。

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][name=""], input[type="radio"]:not([name])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

每个单选按钮组应至少有 2 个单选按钮。

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  });
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert.fail();
});

assert.isAbove(groupKeys.length, 0);
```

你应该至少有两个 `input` 元素，`type` 为 `checkbox`（复选框），它们是 `#survey-form` 的子元素。

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert.isAtLeast(els.length, 2);
```

你在 `#survey-form` 中的所有复选框都应该有 `value` 属性和值。

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll(
  '#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

你至少应该有一个 `textarea` 元素，它是 `#survey-form` 的子元素。

```js
const el = document.querySelector('#survey-form textarea');
assert.isNotNull(el);
```

你应该有一个 `input` 或 `button` 元素，其 `id` 为 `submit`。

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.isTrue(el.tagName === 'INPUT' || el.tagName === 'BUTTON');
```

你的 `#submit` 应该有一个 `type`，值为 `submit`。

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.strictEqual(el.type, 'submit');
```

你的 `#submit` 应该是 `#survey-form` 的子元素。

```js
const el = document.querySelector('#survey-form #submit');
assert.isNotNull(el);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Survey Form</title>
  </head>

  <body></body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="container">
      <header class="header">
        <h1 id="title" class="text-center">freeCodeCamp Survey Form</h1>
        <p id="description" class="description text-center">
          Thank you for taking the time to help us improve the platform
        </p>
      </header>
      <form id="survey-form">
        <div class="form-group">
          <label id="name-label" for="name">Name</label>
          <input
            type="text"
            name="name"
            id="name"
            class="form-control"
            placeholder="Enter your name"
            required
          />
        </div>
        <div class="form-group">
          <label id="email-label" for="email">Email</label>
          <input
            type="email"
            name="email"
            id="email"
            class="form-control"
            placeholder="Enter your Email"
            required
          />
        </div>
        <div class="form-group">
          <label id="number-label" for="number"
            >Age<span class="clue">(optional)</span></label
          >
          <input
            type="number"
            name="age"
            id="number"
            min="10"
            max="99"
            class="form-control"
            placeholder="Age"
          />
        </div>
        <div class="form-group">
          <p>Which option best describes your current role?</p>
          <select id="dropdown" name="role" class="form-control" required>
            <option disabled selected value>Select current role</option>
            <option value="student">Student</option>
            <option value="job">Full Time Job</option>
            <option value="learner">Full Time Learner</option>
            <option value="preferNo">Prefer not to say</option>
            <option value="other">Other</option>
          </select>
        </div>

        <div class="form-group">
          <p>Would you recommend freeCodeCamp to a friend?</p>
          <label>
            <input
              name="user-recommend"
              value="definitely"
              type="radio"
              class="input-radio"
              checked
            />Definitely</label
          >
          <label>
            <input
              name="user-recommend"
              value="maybe"
              type="radio"
              class="input-radio"
            />Maybe</label
          >

          <label
            ><input
              name="user-recommend"
              value="not-sure"
              type="radio"
              class="input-radio"
            />Not sure</label
          >
        </div>

        <div class="form-group">
          <p>What is your favorite feature of freeCodeCamp?</p>
          <select id="most-like" name="mostLike" class="form-control" required>
            <option disabled selected value>Select an option</option>
            <option value="challenges">Challenges</option>
            <option value="projects">Projects</option>
            <option value="community">Community</option>
            <option value="openSource">Open Source</option>
          </select>
        </div>

        <div class="form-group">
          <p>
            What would you like to see improved?
            <span class="clue">(Check all that apply)</span>
          </p>

          <label
            ><input
              name="prefer"
              value="front-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Front-end Projects</label
          >
          <label>
            <input
              name="prefer"
              value="back-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Back-end Projects</label
          >
          <label
            ><input
              name="prefer"
              value="data-visualization"
              type="checkbox"
              class="input-checkbox"
            />Data Visualization</label
          >
          <label
            ><input
              name="prefer"
              value="challenges"
              type="checkbox"
              class="input-checkbox"
            />Challenges</label
          >
          <label
            ><input
              name="prefer"
              value="open-source-community"
              type="checkbox"
              class="input-checkbox"
            />Open Source Community</label
          >
          <label
            ><input
              name="prefer"
              value="gitter-help-rooms"
              type="checkbox"
              class="input-checkbox"
            />Gitter help rooms</label
          >
          <label
            ><input
              name="prefer"
              value="videos"
              type="checkbox"
              class="input-checkbox"
            />Videos</label
          >
          <label
            ><input
              name="prefer"
              value="city-meetups"
              type="checkbox"
              class="input-checkbox"
            />City Meetups</label
          >
          <label
            ><input
              name="prefer"
              value="wiki"
              type="checkbox"
              class="input-checkbox"
            />Wiki</label
          >
          <label
            ><input
              name="prefer"
              value="forum"
              type="checkbox"
              class="input-checkbox"
            />Forum</label
          >
          <label
            ><input
              name="prefer"
              value="additional-courses"
              type="checkbox"
              class="input-checkbox"
            />Additional Courses</label
          >
        </div>

        <div class="form-group">
          <p>Any comments or suggestions?</p>
          <textarea
            id="comments"
            class="input-textarea"
            name="comment"
            placeholder="Enter your comment here..."
          ></textarea>
        </div>

        <div class="form-group">
          <button type="submit" id="submit" class="submit-button">
            Submit
          </button>
        </div>
      </form>
    </div>
  </body>
</html>
```
