---
id: 68e6a99a0fd7386fb66e156f
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop üben Sie, Kontrollkästchen zu gestalten, indem Sie ein Einstellungsfenster für ein Spiel erstellen.

Der gesamte HTML-Boilerplate-Code (`DOCTYPE`, `html`, `head` und `body`) wurde für Sie bereitgestellt.

Erstellen Sie innerhalb des `body` ein `div`-Element mit einem verschachtelten `h1`-Element. Geben Sie dem `h1`-Element den Text `Game Settings`

# --hints--

Sie sollten ein `div`-Element haben.

```js
assert.lengthOf(document.querySelectorAll("div"), 1);
```

Sie sollten ein `h1`-Element haben.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

Ihr `h1`-Element sollte sich innerhalb Ihres `div`-Elements befinden.

```js
assert.exists(document.querySelector("div")?.querySelector("h1"));
```

Ihr `h1`-Element sollte den Text `Game Settings` enthalten.

```js
assert.match(document.querySelector("h1")?.textContent, /Game Settings/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
