---
id: 68e9a6201e2efa142d62bfe1
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Als Nächstes vergrößern Sie die Kontrollkästchen für eine bessere Sichtbarkeit.

Beginnen Sie damit, einen Selektor für `input` einzurichten, der speziell Ihr `[type="checkbox"]` anspricht.

Innerhalb Ihres Selektors setzen Sie die `width` und die `height` auf `20px`. Dadurch wird es größer als zuvor.

Und um mit Ihrer `cursor`-Einstellung übereinzustimmen, die bei den Bezeichnungen gesetzt wurde, fügen Sie `cursor` hinzu und weisen ihm den Wert `pointer` zu. Danach wird beim Überfahren der Kontrollkästchen ein Zeiger angezeigt.

# --hints--

Sie sollten einen Selektor in Ihrem CSS von `input[type="checkbox"]` haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]'));
```

Sie sollten eine `width`-Eigenschaft zu Ihrem `input[type="checkbox"]`-Selektor hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width);
```

Sie sollten eine `width`-Eigenschaft mit dem Wert `20px` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width, "20px");
```

Sie sollten eine `height`-Eigenschaft zu Ihrem `input[type="checkbox"]`-Selektor hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height);
```

Sie sollten eine `height`-Eigenschaft mit dem Wert `20px` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height, "20px");
```

Sie sollten eine `cursor`-Eigenschaft zu Ihrem `input[type="checkbox"]`-Selektor hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor);
```

Sie sollten eine `cursor`-Eigenschaft mit dem Wert `pointer` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
