---
id: 68e9a6bf1a74721a4e43e217
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Jetzt entfernen Sie das standardmäßige Kontrollkästchen, das von Browsern angewendet wird. Fügen Sie innerhalb Ihres `input[type='checkbox']`-Selektors `appearance` mit dem Wert `none` hinzu.

Das Setzen der `appearance`-Eigenschaft auf `none` löscht das Aussehen, das der Browser für Kontrollkästchen anwendet, sodass Sie den gewünschten Stil anzeigen können.

Da das Kontrollkästchen danach nicht mehr sichtbar ist, setzen Sie eine `border` mit einer Dicke von `2px`, einem `solid`-Stil und einem Hex-Code von `#f1be32`.

# --hints--

Sie sollten Ihrem `appearance`-Selektor eine `input[type="checkbox"]`-Eigenschaft hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance);
```

Sie sollten eine `appearance`-Eigenschaft mit dem Wert `none` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance, "none");
```

Sie sollten Ihrem `border`-Selektor eine `input[type="checkbox"]`-Eigenschaft hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border);
```

Sie sollten eine `border`-Eigenschaft mit dem Wert `2px solid #f1be32` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border, "2px solid rgb(241, 190, 50)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
