---
id: 68e9a7bd60c971278d2d4ffc
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

In einer vorherigen Lektion haben Sie Pseudoklassen und Pseudoelemente in CSS kennengelernt. Jetzt wenden Sie dieses Wissen an, indem Sie einen kombinierten Typselektor mit Pseudoklassen-Selektor erstellen.

Beginnen Sie zuerst mit dem Format für den Typselektor, indem Sie ihn als `input[type="checkbox"]` festlegen und `:checked` am Ende anhängen.

Geben Sie ihm anschließend eine `background-color` mit dem Wert `#f1be32`.

Und zuletzt setzen Sie die `border-color` auf den Wert `#e2a60d`.

# --hints--

Sie sollten einen kombinierten Typ- und Pseudoklassen-Selektor von `input[type="checkbox"]:checked` haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked'));
```

Sie sollten eine `background-color`-Eigenschaft zu Ihrem kombinierten Typ- und Pseudoklassen-Selektor von `input[type="checkbox"]:checked` hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor);
```

Sie sollten eine `background-color`-Eigenschaft mit dem Wert `#f1be32` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor, "rgb(241, 190, 50)");
```

Sie sollten eine `border-color`-Eigenschaft zu Ihrem kombinierten Typ- und Pseudoklassen-Selektor von `input[type="checkbox"]:checked` hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor);
```

Sie sollten eine `border-color`-Eigenschaft mit dem Wert `#e2a60d` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor, "rgb(226, 166, 13)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

--fcc-editable-region--

--fcc-editable-region--

```
