---
id: 68e9a73075dc26207320811a
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Da Sie Ihre Kontrollkästchen jetzt wieder sehen können, werden Sie einige Styling-Optionen dafür abschließen. Geben Sie ihnen eine abgerundete Kante, indem Sie Ihrem `border-radius`-Selektor eine `4px` von `input[type="checkbox"]` hinzufügen.

Geben Sie ihm dann eine `background-color` von `white`, damit sich die Mitte Ihres Kontrollkästchens vom Hintergrund des Containers abhebt.

Da Sie eine benutzerdefinierte Transition einrichten, die beim Klicken auf die Kontrollkästchen ausgelöst wird, legen Sie eine `transition` mit dem Wert `all` und `0.3s` fest, damit die Transition sanft über 0,3 Sekunden abläuft und nicht sofort.

# --hints--

Sie sollten Ihrem `border-radius`-Selektor eine `input[type="checkbox"]`-Eigenschaft hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.borderRadius);
```

Sie sollten eine `border-radius`-Eigenschaft mit dem Wert `4px` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.borderRadius, "4px");
```

Sie sollten Ihrem `background-color`-Selektor eine `input[type="checkbox"]`-Eigenschaft hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.backgroundColor);
```

Sie sollten eine `background-color`-Eigenschaft mit dem Wert `white` haben.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.backgroundColor, ["white", "rgb(255,255,255)", "#ffffff"]);
```

Sie sollten Ihrem `transition`-Selektor eine `input[type="checkbox"]`-Eigenschaft hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.transition);
```

Sie sollten `transition` mit dem Wert `all 0.3s` haben.

```js
/**
 * Because CSS is simply the most lovely thing to ever exist, the behaviour of the `all` shorthand does wonky things and will not show up
 * in the computed stylesheet. Instead, we need to query the actual style element and extract the selector, checking the properties as written
 * by the camper.
 * I really hate CSS.
 * <3 Naomi
 */
const { props } = document.querySelector('.fcc-injected-styles')?.innerText?.match(/input\[type=('|")checkbox\1\]\s*\{(?<props>.*)\}/s)?.groups;
assert.match(props, /transition\s*\:\s*all\s+0.3\s*/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
