---
id: 68e9a834dfef752dc35f6f90
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

Da Ihre Kontrollkästchen jetzt die Farbe ändern, wenn der Benutzer darauf klickt, werden Sie ein Styling anwenden, um deutlicher zu machen, dass das Kontrollkästchen ausgewählt wurde.

CSS-Pseudoklassen können verkettet werden. Sie können eine Pseudoklasse oder ein Pseudoelement nach dem anderen hinzufügen, um Elemente anzusprechen, die mehrere Bedingungen erfüllen.

Erstellen Sie einen neuen Selektor für `input[type="checkbox"]`. Fügen Sie die `:checked`-Pseudoklasse gefolgt vom `::after`-Pseudoelement hinzu. So können Sie einen visuellen Hinweis hinzufügen, wenn das Kontrollkästchen aktiviert ist.

Fügen Sie anschließend `content` mit dem Wert `"✓"` hinzu.

# --hints--

Sie sollten einen kombinierten Typselektor mit einem Pseudoklassen-Selektor und Pseudoelement-Selektor von `input[type="checkbox"]:checked::after` haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after'));
```

Sie sollten Ihrem kombinierten Typselektor mit Pseudoklassen-Selektor und Pseudoelement-Selektor von `content` eine `input[type="checkbox"]:checked::after`-Eigenschaft hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content);
```

Sie sollten eine `content`-Eigenschaft mit dem Wert `"✓"` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content, '"✓"');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

--fcc-editable-region--

--fcc-editable-region--

```

