---
id: 68e9a3194160f8f5620ddf21
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Unmittelbar nach Ihrem `h1`-Element erstellen Sie vier `label`-Elemente.

Innerhalb jedes der `label`-Elemente erstellen Sie ein `input`-Element mit dem `type`-Attribut, das auf `"checkbox"` gesetzt ist.

Nach jedem `input`-Element geben Sie Ihren `label`-Elementen nacheinander die folgenden Texte: `Sound Effects`, `Background Music`, `Hard Mode` und `Haptic Feedback`.

# --hints--

Sie sollten vier `label`-Elemente haben.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

Ihr erstes `label`-Element sollte ein `input`-Element enthalten.

```js
assert.exists(document.querySelectorAll("label")[0]?.querySelector("input"));
```

Ihr erstes `input`-Element sollte das `type`-Attribut mit dem Wert `checkbox` haben.

```js
assert.strictEqual(document.querySelectorAll("label")[0]?.querySelector("input")?.type, "checkbox");
```

Ihr erstes `label`-Element sollte den Text `Sound Effects` enthalten.

```js
assert.match(document.querySelectorAll("label")[0]?.textContent, /Sound Effects/i);
```

Ihr `Sound Effects`-Text sollte nach dem `input`-Element stehen.

```js
assert.match(document.querySelectorAll("label")[0]?.innerHTML, /\<input[^>]*>\s*Sound Effects/i)
```

Ihr zweites `label`-Element sollte ein `input`-Element enthalten.

```js
assert.exists(document.querySelectorAll("label")[1]?.querySelector("input"));
```

Ihr zweites `input`-Element sollte das `type`-Attribut mit dem Wert `checkbox` haben.

```js
assert.strictEqual(document.querySelectorAll("label")[1]?.querySelector("input")?.type, "checkbox");
```

Ihr zweites `label`-Element sollte den Text `Background Music` enthalten.

```js
assert.match(document.querySelectorAll("label")[1]?.textContent, /Background Music/i)
```

Ihr `Background Music`-Text sollte nach dem `input`-Element stehen.

```js
assert.match(document.querySelectorAll("label")[1]?.innerHTML, /\<input[^>]*>\s*Background Music/i)
```

Ihr drittes `label`-Element sollte ein `input`-Element enthalten.

```js
assert.exists(document.querySelectorAll("label")[2]?.querySelector("input"));
```

Ihr drittes `input`-Element sollte das `type`-Attribut mit dem Wert `checkbox` haben.

```js
assert.strictEqual(document.querySelectorAll("label")[2]?.querySelector("input")?.type, "checkbox");
```

Ihr drittes `label`-Element sollte den Text `Hard Mode` enthalten.

```js
assert.match(document.querySelectorAll("label")[2]?.textContent, /Hard Mode/i)
```

Ihr `Hard Mode`-Text sollte nach dem `input`-Element stehen.

```js
assert.match(document.querySelectorAll("label")[2]?.innerHTML, /\<input[^>]*>\s*Hard Mode/i)
```

Ihr viertes `label`-Element sollte ein `input`-Element enthalten.

```js
assert.exists(document.querySelectorAll("label")[3]?.querySelector("input"));
```

Ihr viertes `input`-Element sollte das `type`-Attribut mit dem Wert `checkbox` haben.

```js
assert.strictEqual(document.querySelectorAll("label")[3]?.querySelector("input")?.type, "checkbox");
```

Ihr viertes `label`-Element sollte den Text `Haptic Feedback` enthalten.

```js
assert.match(document.querySelectorAll("label")[3]?.textContent, /Haptic Feedback/i)
```

Ihr `Haptic Feedback`-Text sollte nach dem `input`-Element stehen.

```js
assert.match(document.querySelectorAll("label")[3]?.innerHTML, /\<input[^>]*>\s*Haptic Feedback/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```
