---
id: 68ee86e3ee07b01626a392cc
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Setzen Sie innerhalb Ihres `.settings-card`-Selektors die `margin`-Eigenschaft auf `auto`.

Das Setzen der `margin`-Eigenschaft auf `auto` passt automatisch die Außenabstände eines Elements an, um den verbleibenden Platz in seinem Container gleichmäßig zu verteilen. Dies wird häufig verwendet, um Block-Elemente horizontal zu zentrieren.

Und zuletzt setzen Sie eine `text-align`-Eigenschaft mit dem Wert `left`. Dadurch wird der inline-Inhalt, wie Text, an der linken Seite seines enthaltenden Elements ausgerichtet.

# --hints--

Sie sollten eine `margin`-Eigenschaft zum `settings-card`-Klassen-Selektor hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin);
```

Sie sollten eine `margin`-Eigenschaft mit dem Wert `auto` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin, "auto");
```

Sie sollten eine `text-align`-Eigenschaft zum `settings-card`-Klassen-Selektor hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign);
```

Sie sollten eine `text-align`-Eigenschaft mit dem Wert `left` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign, "left");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
