---
id: 68f15cd397503429f0345262
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

In der Deklaration für den `input[type="checkbox"]:checked::after`-Selektor setzen Sie die `display` auf `block`. Wenn Sie Ihre `display`-Eigenschaft auf `block` setzen, wird das Element zu einem Block-Element, das heißt, es nimmt die gesamte Breite seines Containers ein und beginnt in einer neuen Zeile, sodass Sie seine Breite, Höhe und Abstände leichter steuern können.

Als Nächstes setzen Sie `text-align` auf `center`. Dadurch wird der inline Inhalt (wie Text oder Inline-Elemente) horizontal innerhalb des Blocks zentriert. Der Block nimmt weiterhin die volle Breite seines Containers ein, aber alles darin wird zentriert ausgerichtet.

# --hints--

Sie sollten eine `display`-Eigenschaft zu Ihrem kombinierten Typselektor mit einem Pseudoklassen-Selektor und Pseudoelement-Selektor von `input[type="checkbox"]:checked::after` hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display);
```

Sie sollten eine `display`-Eigenschaft mit dem Wert `block` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display, "block");
```

Sie sollten eine `text-align`-Eigenschaft zu Ihrem kombinierten Typselektor mit einem Pseudoklassen-Selektor und Pseudoelement-Selektor von `input[type="checkbox"]:checked::after` hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign);
```

Sie sollten eine `text-align`-Eigenschaft mit dem Wert `center` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type="checkbox"]:checked::after {
  content: "✓";
--fcc-editable-region--

--fcc-editable-region--
}
```
