---
id: 68f16192e8a66f3b948d3ebd
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Um den visuellen Effekt Ihres Kontrollkästchens fertig einzurichten, setzen Sie die `font-weight`-Eigenschaft auf den Wert `bold`. Dadurch wird die Sichtbarkeit des Häkchens erhöht.

Da es jetzt leichter zu sehen ist, ändern Sie die `color` auf den Wert `white`. Dies ändert die Farbe des Häkchens im Kontrollkästchen, wenn es aktiviert ist.

Zuletzt ändern Sie die `line-height` auf den Wert `20px`.

Damit haben Sie das Einstellungsfenster für das Spiel abgeschlossen!

# --hints--

Sie sollten eine `font-weight`-Eigenschaft zu Ihrem kombinierten Typselektor mit einem Pseudoklassen-Selektor und Pseudoelement-Selektor von `input[type="checkbox"]:checked::after` hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.fontWeight);
```

Sie sollten eine `font-weight`-Eigenschaft mit dem Wert `bold` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.fontWeight, "bold");
```

Sie sollten eine `color`-Eigenschaft zu Ihrem kombinierten Typselektor mit einem Pseudoklassen-Selektor und Pseudoelement-Selektor von `input[type="checkbox"]:checked::after` hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.color);
```

Sie sollten eine `color`-Eigenschaft mit dem Wert `white` haben.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.color, ["white", "#ffffff", "rgb(255,255,255)"]);
```

Sie sollten eine `line-height`-Eigenschaft zu Ihrem kombinierten Typselektor mit einem Pseudoklassen-Selektor und Pseudoelement-Selektor von `input[type="checkbox"]:checked::after` hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.lineHeight);
```

Sie sollten eine `line-height`-Eigenschaft mit dem Wert `20px` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.lineHeight, "20px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type="checkbox"]:checked::after {
  content: "✓";
  display: block;
  text-align: center;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type='checkbox']:checked:after {
  content: '✓';
  display: block;
  text-align: center;
  font-weight: bold;
  color: white;
  line-height: 20px;
}
```

