---
id: 68f15ef0446464325ce1aa45
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Sie benötigen etwas Abstand zwischen den Kontrollkästchen und den Bezeichnungen. Beginnen Sie damit, in Ihrem CSS den Selektor für `label` zu verwenden und eine `display`-Eigenschaft mit dem Wert `block` zu setzen.

Als Nächstes setzen Sie eine `margin`-Eigenschaft mit dem Wert `8px auto`, um vertikalen Abstand zwischen den Elementen hinzuzufügen. Dadurch entsteht ein gleichmäßiger Abstand oberhalb und unterhalb jedes Elements, was hilft, den Inhalt zu trennen und die Lesbarkeit zu verbessern.

Zum Schluss setzen Sie die `cursor`-Eigenschaft auf `pointer` bei den `label`-Elementen. Dadurch ändert sich der Cursor beim Überfahren eines `label` in ein Hand-Symbol, was signalisiert, dass das Element anklickbar ist, und die Benutzererfahrung insgesamt verbessert.

# --hints--

Sie sollten einen Typselektor in Ihrem CSS für `label` haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("label"));
```

Sie sollten eine `display`-Eigenschaft zu Ihrem `label`-Selektor hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.display);
```

Sie sollten eine `display`-Eigenschaft mit dem Wert `block` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.display, "block");
```

Sie sollten die `margin`-Eigenschaft zu Ihrem `label`-Selektor hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.margin);
```

Sie sollten eine `margin`-Eigenschaft mit dem Wert `8px auto` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("label")?.margin, "8px auto");
```

Sie sollten eine `cursor`-Eigenschaft zu Ihrem `label`-Selektor hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.cursor);
```

Sie sollten eine `cursor`-Eigenschaft mit dem Wert `pointer` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}
--fcc-editable-region--

--fcc-editable-region--
```
