---
id: 68f1560aff1214103ab6b59b
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Jetzt beginnen Sie damit, die Seite mit etwas CSS-Styling aufzuhübschen. Beginnen Sie damit, den `body`-Selektor zu erstellen.

Setzen Sie die `body` so, dass sie eine `height`-Eigenschaft mit dem Wert `100vh` und eine `background-color`-Eigenschaft mit dem Wert `#f0f0f0` hat.

Die `height` von `100vh` sorgt dafür, dass die `body` die volle Höhe des Browser-Viewports einnimmt, während die hellgraue Hintergrundfarbe eine dezente Basis für die Seite bietet.

Legen Sie zuletzt eine `text-align`-Eigenschaft mit dem Wert `center` fest. Dadurch wird sämtlicher inline-Inhalt auf der Seite zentriert, sofern ein Kindelement diese Ausrichtung nicht mit seiner eigenen überschreibt.

# --hints--

Sie sollten den `body`-Typselektor verwenden.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

Sie sollten die `height`-Eigenschaft innerhalb des `body`-Typselektors hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.height);
```

Sie sollten eine `height`-Eigenschaft mit dem Wert `100vh` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.height, "100vh");
```

Sie sollten die `background-color`-Eigenschaft innerhalb des `body`-Typselektors hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor);
```

Sie sollten eine `background-color`-Eigenschaft mit dem Wert `#f0f0f0` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor, "rgb(240, 240, 240)");
```

Sie sollten die `text-align`-Eigenschaft innerhalb des `body`-Typselektors hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.textAlign);
```

Sie sollten eine `text-align`-Eigenschaft mit dem Wert `center` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
