---
id: 68f29753fd5524530c6e65c8
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Als Nächstes setzen Sie eine `vertical-align`-Eigenschaft mit dem Wert `middle`.

Die `vertical-align`-Eigenschaft steuert, wie `inline`- oder `inline-block`-Elemente vertikal mit dem umgebenden Text oder anderen Inline-Elementen ausgerichtet werden. Sie wird häufig verwendet, um die vertikale Position von Elementen wie Bildern, Symbolen oder Text innerhalb einer Zeile anzupassen.

Schließen Sie dann Ihr Kontrollkästchen mit einer `margin`-Eigenschaft von `15px` ab.

# --hints--

Sie sollten Ihrem `vertical-align`-Selektor eine `input[type="checkbox"]`-Eigenschaft hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.verticalAlign);
```

Sie sollten eine `vertical-align`-Eigenschaft mit dem Wert `middle` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.verticalAlign, "middle");
```

Sie sollten Ihrem `margin-right`-Selektor eine `input[type="checkbox"]`-Eigenschaft hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.marginRight);
```

Sie sollten eine `margin-right`-Eigenschaft mit dem Wert `15px` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.marginRight, "15px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
