---
id: 68f29dbf9125476b4808d6fb
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Jetzt ist es Zeit, den `settings-card`-Container zu gestalten. Erstellen Sie einen Klassenselektor für `settings-card`.

Hier wird die gesamte Formatierung für den Container erfolgen. Setzen Sie die `max-width` auf `250px`, um die Gesamtgröße Ihres Containers festzulegen.

Als Nächstes setzen Sie Ihre `padding` auf `20px`, damit Ihr Inhalt Abstand zum Rahmen des Containers hat.

Danach erstellen Sie eine abgerundete Kante, indem Sie Ihre `border-radius` auf `10px` setzen.

Anschließend setzen Sie eine `box-shadow` mit den Werten `0 2px 6px rgba(0,0,0,0.2)`. Dies erzeugt einen dezenten „schwebenden“ Effekt, der dem Container Tiefe verleiht.

# --hints--

Sie sollten einen Klassenselektor in Ihrem CSS für `settings-card` haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".settings-card"));
```

Sie sollten die `max-width`-Eigenschaft zum `.settings-card`-Klassenselektor hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth);
```

Sie sollten eine `max-width`-Eigenschaft mit dem Wert `250px` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth, "250px");
```

Sie sollten die `padding`-Eigenschaft zum `.settings-card`-Klassenselektor hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding);
```

Sie sollten eine `padding`-Eigenschaft mit dem Wert `20px` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding, "20px");
```

Sie sollten die `border-radius`-Eigenschaft zu Ihrem `.settings-card`-Klassenselektor hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius);
```

Sie sollten eine `border-radius`-Eigenschaft mit dem Wert `10px` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius, "10px");
```

Sie sollten die `box-shadow`-Eigenschaft zum `.settings-card`-Klassenselektor hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow);
```

Sie sollten eine `box-shadow` mit dem Wert `0 2px 6px rgba(0,0,0,0.2)` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow, "rgba(0, 0, 0, 0.2) 0px 2px 6px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--
```
