---
id: 68dbc6b7e8f794a2ae69bc74
title: Реалізувати алгоритм N-ферзів
challengeType: 27
dashedName: implement-the-n-queens-algorithm
---

# --description--

Задача N-ферзів полягає в тому, щоб розмістити N ферзів на шахівниці розміром N×N так, щоб жодні два ферзі не нападали один на одного (жодні два не знаходилися в одному рядку, стовпці або діагоналі).

Наприклад, якщо є дошка 4x4, одним із правильних розташувань є:

```md
[1, 3, 0, 2]
```

Це означає, що в рядку 0 ферзь розміщений у стовпці 1; у рядку 1 ферзь розміщений у стовпці 3; у рядку 2 ферзь розміщений у стовпці 0; а в рядку 3 ферзь розміщений у стовпці 2.

Візуально це розташування виглядає так:

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

Де `Q` позначає ферзя, а `.` — порожню клітинку.

У цій лабораторній роботі ви реалізуєте розв’язувач задачі N-ферзів, використовуючи підхід пошуку у глибину.


**Мета**: Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути функція з назвою `dfs_n_queens`.
2. Функція має приймати рівно один аргумент: ціле число `n`.
3. Якщо `n` менше за `1`, функція має повертати порожній список (`[]`).
4. Функція має повертати список розв’язків; кожен розв’язок — це список довжини `n`, де елемент з індексом `i` — це індекс стовпця (починаючи з 0) ферзя в рядку `i`.

# --hints--

У вас має бути функція з назвою `dfs_n_queens`, яка приймає один аргумент.

```js
({ test: () => runPython(`
  import inspect
  assert inspect.isfunction(dfs_n_queens)
  sig = inspect.signature(dfs_n_queens)
  assert len(sig.parameters) == 1
`) })
```

Якщо `n` менше за `1`, функція має повертати порожній список.

```js
({ test: () => runPython(`
  assert dfs_n_queens(0) == []
  assert dfs_n_queens(-1) == []
  assert dfs_n_queens(-5) == []
`) })
```

Функція має повертати список розв’язків, де кожен розв’язок — це список довжини `n`.

```js
({ test: () => runPython(`
  result = dfs_n_queens(4)
  assert isinstance(result, list)
  for solution in result:
    assert isinstance(solution, list)
    assert len(solution) == 4
`) })
```

`dfs_n_queens(1)` має повертати `[[0]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(1) == [[0]]
`) })
```

`dfs_n_queens(2)` має повертати `[]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(2) == []
`) })
```

`dfs_n_queens(3)` має повертати `[]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(3) == []
`) })
```

`dfs_n_queens(4)` має повертати `[[1, 3, 0, 2], [2, 0, 3, 1]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(4) == [[1, 3, 0, 2], [2, 0, 3, 1]]
`) })
```

`dfs_n_queens(5)` має повертати `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(5) == [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]
`) })
```

`len(dfs_n_queens(5))` має бути `10`.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(5)) == 10
`) })
```

`len(dfs_n_queens(8))` має бути `92`.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(8)) == 92
`) })
```

`dfs_n_queens` має повертати правильний результат.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(7)) == 40
  assert dfs_n_queens(6) == [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def dfs_n_queens(n):
    if n < 1:
        return []

    results = []

    def is_safe(queens, row, col):
        for r, c in enumerate(queens):
            if c == col or abs(row - r) == abs(col - c):
                return False
        return True

    def dfs(queens):
        row = len(queens)
        if row == n:
            results.append(queens[:])
            return
        for col in range(n):
            if is_safe(queens, row, col):
                queens.append(col)
                dfs(queens)
                queens.pop()

    dfs([])
    return results
```
