---
id: 0095d0567b6817114989a1dd
title: 第一步驟
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將使用 JavaScript 實作最短路徑演算法。你會撰寫一個函式，計算加權圖中節點之間的最短路徑，並傳回距離和所走的路徑。

例如，給定一個城市由不同距離的道路連接的圖，演算法會找出從一個城市到另一個城市的最短路線。如果你想從城市 A 旅行到城市 D，演算法可能會發現走 A ⇨ B ⇨ C ⇨ D（總共 15 公里）比直接走 A ⇨ D（20 公里）更短。

要開始，定義一個名為 `INF` 的變數，並將其指定為 `Infinity`。之後，你會用它來顯示兩個節點之間沒有直接連線。

# --hints--

你應該定義一個名為 `INF` 的變數。

```js
assert.isDefined(INF);
```

你的 `INF` 變數應該有 `Infinity` 的值。

```js
assert.strictEqual(INF, Infinity);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
