---
id: 69c918841987bf16c1d1bacb
title: 第 20 步驟
challengeType: 1
dashedName: step-20
---

# --description--

當主要迴圈結束後，你需要決定要顯示哪個節點的結果。

如果有提供特定的 `targetNode`，只顯示該節點的結果。否則，顯示所有節點的結果。

在 JavaScript 中，你可以使用 `[...Array(n).keys()]` 取得陣列所有索引作為可迭代物件。

在外層 `for` 迴圈之後，創建一個名為 `targets` 的變數。使用三元運算式，如果 `targetNode` 嚴格不等於 `null`，就指定 `[targetNode]`，否則指定 `[...Array(n).keys()]`。

# --hints--

你應該在主要迴圈之後創建一個 `targets` 變數。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+targets\s*=/
);
```

你的 `targets` 變數應該使用三元運算式檢查 `targetNode !== null`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /targets\s*=\s*targetNode\s*!==\s*null\s*\?/
);
```

當有提供標的時，`targets` 應該是 `[targetNode]`，否則應該是 `[...Array(n).keys()]`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /\?\s*\[\s*targetNode\s*\]\s*:\s*\[\s*\.\.\.\s*Array\s*\(\s*n\s*\)\s*\.keys\s*\(\s*\)\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
