---
id: 69c918841987bf16c1d1bac7
title: 第 16 步
challengeType: 1
dashedName: step-16
---

# --description--

現在當前節點已標記為已訪問，你需要查看它所有的鄰居，看看是否能透過當前節點更高效地到達它們。

在 `visited[current] = true` 之後，新增一個鄰居 `for` 迴圈，迭代從 `nodeNo` 到 `0` 的每個 `n`。請注意，這是與第一個迴圈同層級的另一個內部迴圈。

在迴圈內，宣告一個名為 `distance` 的變數並指定 `matrix[current][nodeNo]` 給它。這會給你當前節點與其鄰居之間的邊權重。

# --hints--

你應該在標記當前節點為已訪問後，新增一個鄰居 `for` 迴圈。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(/
);
```

鄰居 `for` 迴圈應該從 `0` 迭代到 `n`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(\s*(var|let)\s+\w+\s*=\s*0\s*;\s*\w+\s*<\s*n\s*;\s*\w+\+\+\s*\)/
);
```

你應該在鄰居迴圈內有一個 `distance` 變數，並指定 `matrix[current][nodeNo]` 給它。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distance\s*=\s*matrix\s*\[\s*current\s*\]\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
