---
id: 69c918841987bf16c1d1babd
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

你需要追蹤從起始節點到每個其他節點的最短已知距離。開始時，你會假設每個節點的距離都是無限遠。

在 JavaScript 中，你可以使用 `new Array(n).fill(value)` 創建一個預先填滿特定值的陣列。舉例來說，`new Array(3).fill(0)` 會創建 `[0, 0, 0]`。

在 `shortestPath` 函式（程式）內，創建一個名為 `distances` 的變數，並將其指定為 `new Array(n).fill(INF)`。

# --hints--

你應該在 `distances` 內創建一個名為 `shortestPath` 的變數。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distances\s*=/
);
```

你的 `distances` 變數應該被設值為 `new Array(n).fill(INF)`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*INF\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;

}
--fcc-editable-region--
```
