---
id: 69c918841987bf16c1d1bac4
title: 第 13 步
challengeType: 1
dashedName: step-13
---

# --description--

如果條件為真，表示目前節點是迄今為止找到的最佳未造訪候選節點。你需要更新追蹤變數以反映這一點。

在 `if` 述語內，將 `minDistance` 更新為 `distances[nodeNo]`，並將 `current` 更新為 `nodeNo`。

# --hints--

你應該在 `minDistance` 區塊內將 `distances[nodeNo]` 設為 `if`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /minDistance\s*=\s*distances\s*\[\s*nodeNo\s*\]/
);
```

你應該在 `current` 區塊內將 `nodeNo` 設為 `if`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /current\s*=\s*nodeNo/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        --fcc-editable-region--

        --fcc-editable-region--
      }
    }
  }
}
```
