---
id: 69c918841987bf16c1d1babb
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

現在你將創建執行所有工作的主要函式。

創建一個名為 `shortestPath` 的函式，該函式接受三個參數：`matrix`、`startNode` 和 `targetNode`。

在接下來的幾個步驟中，你將為此函式建置邏輯。

# --hints--

你應該創建一個名為 `shortestPath` 的函式。

```js
assert.isFunction(shortestPath);
```

你的 `shortestPath` 函式應該有參數 `matrix`、`startNode` 和 `targetNode`。

```js
const params = __helpers.getFunctionParams(shortestPath.toString());
const names = params.map(p => p.name);
assert.include(names, 'matrix');
assert.include(names, 'startNode');
assert.include(names, 'targetNode');
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--

--fcc-editable-region--
```
