---
id: 69c918841987bf16c1d1babf
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

你還需要追蹤到達每個節點的實際路徑。你將儲存一個路徑列表，其中每個項目都是一個節點索引的陣列，代表所走過的路線。

一開始，每個節點的路徑只會包含它自己。你可以使用 `Array.from()` 來創建這個結構：

```js
Array.from({ length: n }, (_, i) => [i]);
```

這會創建一個有 `n` 個元素的陣列，其中索引為 `i` 的元素是 `[i]`。

仍在 `shortestPath` 函式內，創建一個名為 `paths` 的變數並將 `Array.from({ length: n }, (_, i) => [i])` 指定給它。

# --hints--

你應該在 `paths` 函式內創建一個變數 `shortestPath`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+paths\s*=/
);
```

你的 `paths` 變數應該使用 `Array.from` 來初始化。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*=\s*Array\.from/
);
```

你的 `paths` 變數應該使用 `{ length: n }` 作為 `Array.from` 的第一個引數。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\(\s*\{\s*length\s*:\s*n\s*\}/
);
```

`Array.from` 的第二個引數應該是一個回呼 `(_, i) => [i]`，其中 `_` 是未使用元素的佔位符，`i` 是索引。它應該傳回 `[i]`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\([\s\S]*?,\s*function\s*\(\s*\w+\s*,\s*(\w+)\s*\)\s*\{[\s\S]*?return\s*\[\s*\1\s*\]|Array\.from\s*\([\s\S]*?,\s*\(\s*\w+\s*,\s*\w+\s*\)\s*=>\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
