---
id: 69c937c73f23b1140bcd297a
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

`targetNode` 參數是選用的。當未提供時，函式會計算從 `startNode` 到圖中所有其他節點的最短路徑。

為 `targetNode` 參數設定預設值 `null`。

# --hints--

你的 `targetNode` 函式的 `shortestPath` 參數應該預設為 `null`。

```js
assert.match(
  __helpers.removeJSComments(code),
  /targetNode\s*=\s*null/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode) {

}
--fcc-editable-region--
```
