---
id: 69c918841987bf16c1d1bac2
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

在迴圈內，你需要掃描每個節點，找出尚未被造訪且已知距離最小的節點。

在主迴圈內添加一個 `for` 迴圈，該迴圈會迭代每個節點。使用 `nodeNo` 作為迴圈變數，從 `0` 到 `n`。

目前先將迴圈主體留空。

# --hints--

你應該在主迴圈內有第二個（內層）`for` 迴圈，並以 `nodeNo` 作為其迴圈變數。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+nodeNo\s*=\s*0\s*;\s*nodeNo\s*<\s*n\s*;\s*nodeNo\+\+\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

--fcc-editable-region--
  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;

  }
--fcc-editable-region--
}
```
