---
id: 69c918841987bf16c1d1bac0
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

當演算法執行時，你需要追蹤哪些節點已經被處理過，避免重複拜訪它們。

創建一個名為 `visited` 的變數，並將 `new Array(n).fill(false)` 指定給它。

# --hints--

你應該創建一個名為 `visited` 的變數。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+visited\s*=/
);
```

你的 `visited` 變數應該設值為 `new Array(n).fill(false)`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*false\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
--fcc-editable-region--

--fcc-editable-region--
}
```
