---
id: 69c918841987bf16c1d1bac5
title: 第 14 步
challengeType: 1
dashedName: step-14
---

# --description--

在掃描所有節點以尋找最近且未造訪的節點後，你需要處理找不到有效節點的情況。這種情況發生在所有剩餘節點都無法到達時。

在內部 `for` 迴圈之後，但仍在外部迴圈內，新增一個 `if` 述語，檢查 `current` 是否嚴格相等於 `-1`。如果為真，請在 `break` 區塊內使用 `if` 來提前停止外部迴圈。

# --hints--

你應該新增一個 `if` 述語來檢查 `current === -1`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)/
);
```

你應該在 `break` 區塊內使用 `if`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)[\s\S]*?break/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
