---
id: 69c918841987bf16c1d1babe
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

從起始節點到它本身的距離永遠是 `0`。你必須更新 `distances` 陣列以反映這一點。

在 `distances` 陣列之後，將 `distances[startNode]` 設為 `0`。

# --hints--

你應該將 `distances[startNode]` 設為 `0`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /distances\s*\[\s*startNode\s*\]\s*=\s*0/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);

}
--fcc-editable-region--
```
