---
id: 69c918841987bf16c1d1baba
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

你需要一個二維陣列來表述圖的鄰接矩陣。每個值代表兩個節點之間邊的權重（距離）。值為 `INF` 表示沒有直接的邊。

創建另一個名為 `adjMatrix` 的變數，並指定給它一個包含以下值的二維陣列：

```js
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

你應該創建一個名為 `adjMatrix` 的變數。

```js
assert.isDefined(adjMatrix);
```

`adjMatrix` 應該是一個有 6 列 6 行的二維陣列。

```js
assert.isArray(adjMatrix);
assert.lengthOf(adjMatrix, 6);
adjMatrix.forEach(row => assert.lengthOf(row, 6));
```

`adjMatrix` 應該包含正確的值。

```js
assert.deepEqual(adjMatrix[0], [0, 5, 3, INF, 11, INF]);
assert.deepEqual(adjMatrix[1], [5, 0, 1, INF, INF, 2]);
assert.deepEqual(adjMatrix[2], [3, 1, 0, 1, 5, INF]);
assert.deepEqual(adjMatrix[3], [INF, INF, 1, 0, 9, 3]);
assert.deepEqual(adjMatrix[4], [11, INF, 5, 9, 0, INF]);
assert.deepEqual(adjMatrix[5], [INF, 2, INF, 3, INF, 0]);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
--fcc-editable-region--

--fcc-editable-region--
```
