---
id: 69c918841987bf16c1d1bac3
title: 第 12 步
challengeType: 1
dashedName: step-12
---

# --description--

在內層 `for` 迴圈中，你需要檢查目前節點是否未被造訪且比你目前找到的最佳節點更接近。

新增一個 `if` 述語，檢查 `nodeNo` 是否尚未造訪，且 `distances[nodeNo]` 是否小於 `minDistance`。

目前先將主體留空。

# --hints--

你應該在內層 `if` 迴圈中有一個條件為 `!visited[nodeNo] && distances[nodeNo] < minDistance` 的 `for` 述語。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*!visited\s*\[\s*nodeNo\s*\]\s*&&\s*distances\s*\[\s*nodeNo\s*\]\s*<\s*minDistance\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
