---
id: 69c918841987bf16c1d1bac8
title: 第 17 步
challengeType: 1
dashedName: step-17
---

# --description--

在更新距離之前，你需要確認鄰居是否值得考慮。必須有實際的邊存在（`distance !== INF`），且該鄰居尚未被造訪。

在鄰居 `for` 的迴圈內，新增一個 `if` 述語，檢查 `distance` 是否不嚴格相等於 `INF` **且** `nodeNo` 尚未被造訪。

在 `if` 區塊內，宣告一個變數 `newDistance`，並指定其值為 `distances[current] + distance`。這是透過目前節點到達該鄰居的總成本。

# --hints--

你應該新增一個 `if` 述語，條件為 `distance !== INF && !visited[nodeNo]`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*distance\s*!==\s*INF\s*&&\s*!\s*visited\s*\[\s*\w+\s*\]\s*\)/
);
```

你應該在 `newDistance` 區塊內宣告一個 `distances[current] + distance`，其值為 `if`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+newDistance\s*=\s*distances\s*\[\s*current\s*\]\s*\+\s*distance/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
