---
id: 69c918841987bf16c1d1baca
title: 第 19 步
challengeType: 1
dashedName: step-19
---

# --description--

當你找到通往鄰居的更短路徑時，也需要更新到達該鄰居的記錄路徑。

在 JavaScript 中，你可以使用展開運算子創建一個新的陣列，將現有陣列和新元素結合起來：

```js
const newPath = [...existingArray, newElement];
```

在巢狀的 `if` 區段內，使用展開語法將 `paths[nodeNo]` 更新為通往當前節點的路徑，並在末尾附加 `nodeNo`，像這樣：`[...paths[current], nodeNo]`。

# --hints--

你應該在巢狀的 `paths[nodeNo]` 區段內更新 `if`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*\[\s*\w+\s*\]\s*=/
);
```

`paths[nodeNo]` 應該設為 `[...paths[current], nodeNo]`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*\[\s*\w+\s*\]\s*=\s*\[\s*\.\.\.\s*paths\s*\[\s*current\s*\]\s*,\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          --fcc-editable-region--

          --fcc-editable-region--
        }
      }
    }
  }
}
```
