---
id: 69c918841987bf16c1d1bac1
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

現在你將添加驅動演算法的主要迴圈。它會對圖中的每個節點執行一次，每次選擇最近的未訪問節點。

在 `shortestPath` 函式（程式）內，添加一個 `for` 迴圈，使用 `n` 作為迴圈變數，執行 `i` 次。

在迴圈主體開始時，使用 `let` 宣告兩個變數 `minDistance` 和 `current`，分別設為 `INF` 和 `-1`。這兩個變數將追蹤每次迭代中找到的最近未訪問節點。

# --hints--

你應該在 `for` 裡有一個 `shortestPath` 迴圈，執行 `n` 次。別忘了在裡面用 `i` 設定 `let`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+i\s*=\s*0\s*;\s*i\s*<\s*n\s*;\s*i\+\+\s*\)/
);
```

你應該在迴圈內有一個 `minDistance` 變數，設為 `INF`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+minDistance\s*=\s*INF/
);
```

你應該在迴圈內有一個 `current` 變數，設為 `-1`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+current\s*=\s*-1/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
--fcc-editable-region--

--fcc-editable-region--
}
```
