---
id: 5f332a88dc25a0fd25c7687a
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

The name of the cafe is `CAMPER CAFE`. Add an `h1` element within your `main` element. Give it the name of the cafe in capitalized letters to make it stand out.

# --hints--

يجب أن يكون لديك العلامة الأفتاح (opening tag) الآتية `<h1>`.

```js
assert(code.match(/<h1>/i));
```

يجب أن يكون لديك العلامة الإغلاق (closing tag) الآتي `</h1>`.

```js
assert(code.match(/<\/h1>/i));
```

لا يجب عليك تغيير عنصر `main` الحالي.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

يجب أن يكون عنصرك `h1` داخل عنصرك `main`.

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

يجب أن يحتوي العنصر `h1` على النص `CAMPER CAFE`.

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

