---
id: 594810f028c0303b75339ad8
title: Matriksi ya zig-zag
challengeType: 1
forumTopicId: 302348
dashedName: zig-zag-matrix
---

# --description--

Safu ya data ya 'zig-zag' ni mpangilio wa mraba wa nambari kamili za kwanza $N^2$, ambapo nambari huongezeka mfululizo unapoenda kwa njia ya zig-zag kwenye miondoko ya anti-diagonal ya safu ya data.

Kwa mfano, kwa ingizo `5`, matokeo yafuatayo yanapaswa kuzalishwa:

<pre>
 0  1  5  6 14
 2  4  7 13 15
 3  8 12 16 21
 9 11 17 20 22
10 18 19 23 24
</pre>

# --instructions--

Andika kitendakazi kinachopokea ukubwa wa matriksi ya zig-zag, na kurudisha matriksi inayolingana kama safu ya data ya pande mbili.

# --before-each--

```js
const zm1 = [[0]];
const zm2 = [[0, 1], [2, 3]];
const zm5 = [
  [0, 1, 5, 6, 14],
  [2, 4, 7, 13, 15],
  [3, 8, 12, 16, 21],
  [9, 11, 17, 20, 22],
  [10, 18, 19, 23, 24]
];
```

# --hints--

ZigZagMatrix inapaswa kuwa kitendakazi.

```js
assert.equal(typeof ZigZagMatrix, 'function');
```

ZigZagMatrix inapaswa kurudisha safu ya data.

```js
assert.equal(typeof ZigZagMatrix(1), 'object');
```

ZigZagMatrix inapaswa kurudisha safu ya data ya safu za data zilizopangwa ndani.

```js
assert.equal(typeof ZigZagMatrix(1)[0], 'object');
```

ZigZagMatrix(1) inapaswa kurudisha \[[0]].

```js
assert.deepEqual(ZigZagMatrix(1), zm1);
```

ZigZagMatrix(2) inapaswa kurudisha \[[0, 1], [2, 3]].

```js
assert.deepEqual(ZigZagMatrix(2), zm2);
```

ZigZagMatrix(5) inapaswa kurudisha matriksi iliyotajwa.

```js
assert.deepEqual(ZigZagMatrix(5), zm5);
```

# --seed--

## --seed-contents--

```js
function ZigZagMatrix(n) {

  return [[], []];
}
```

# --solutions--

```js
function ZigZagMatrix(n) {
  const mtx = [];
  for (let i = 0; i < n; i++) {
    mtx[i] = [];
  }

  let i = 1;
  let j = 1;
  for (let e = 0; e < n * n; e++) {
    mtx[i - 1][j - 1] = e;
    if ((i + j) % 2 === 0) {
      // Even stripes
      if (j < n) j++;
      else i += 2;
      if (i > 1) i--;
    } else {
      // Odd stripes
      if (i < n) i++;
      else j += 2;
      if (j > 1) j--;
    }
  }
  return mtx;
}
```
