---
id: 594810f028c0303b75339acc
title: Tatizo la ABC
challengeType: 1
forumTopicId: 302220
dashedName: abc-problem
---

# --description--

Umepatiwa mkusanyiko wa vibonye vya ABC (mfano, vibonye vya alfabeti vya utotoni). Kuna vibonye 20 vyenye herufi mbili kila kibonye. Alfabeti kamili inahakikishwa katika pande zote za vibonye. Mkusanyiko wa mfano wa vibonye:

<pre>(B O)
(X K)
(D Q)
(C P)
(N A)
(G T)
(R E)
(T G)
(Q D)
(F S)
(J W)
(H U)
(V I)
(A N)
(O B)
(E R)
(F S)
(L Y)
(P C)
(Z M)
</pre>

# --instructions--

Tekeleza kitendakazi kinachopokea mfuatano wa herufi (neno) na kuamua kama neno linaweza kuandikwa kwa kutumia mkusanyiko wa vibonye uliopewa.

Miongozo ya kuzingatia:

<ul>
  <li>Mara herufi moja kwenye kibonye itakapotumika, kibonye hicho hakiwezi kutumika tena.</li>
  <li>Kitendakazi kinapaswa kutotegemea herufi kubwa au ndogo.</li>
</ul>

# --before-each--

```js
const words = ['bark', 'BooK', 'TReAT', 'COMMON', 'squAD', 'conFUSE'];
```

# --hints--

`canMakeWord` inapaswa kuwa kitendakazi.

```js
assert(typeof canMakeWord === 'function');
```

`canMakeWord` inapaswa kurudisha thamani ya kweli au siyo kweli.

```js
assert(typeof canMakeWord('hi') === 'boolean');
```

`canMakeWord("bark")` inapaswa kurudisha true.

```js
assert(canMakeWord(words[0]));
```

`canMakeWord("BooK")` inapaswa kurudisha false.

```js
assert(!canMakeWord(words[1]));
```

`canMakeWord("TReAT")` inapaswa kurudisha true.

```js
assert(canMakeWord(words[2]));
```

`canMakeWord("COMMON")` inapaswa kurudisha false.

```js
assert(!canMakeWord(words[3]));
```

`canMakeWord("squAD")` inapaswa kurudisha true.

```js
assert(canMakeWord(words[4]));
```

`canMakeWord("conFUSE")` inapaswa kurudisha true.

```js
assert(canMakeWord(words[5]));
```

# --seed--

## --seed-contents--

```js
function canMakeWord(word) {

}
```

# --solutions--

```js
function canMakeWord(word) {
  const characters = 'BO XK DQ CP NA GT RE TG QD FS JW HU VI AN OB ER FS LY PC ZM';
  const blocks = characters.split(' ').map(pair => pair.split(''));

  const letters = [...word.toUpperCase()];
  let length = letters.length;
  const copy = new Set(blocks);

  letters.forEach(letter => {
    for (let block of copy) {
      const index = block.indexOf(letter);

      if (index !== -1) {
        length--;
        copy.delete(block);
        break;
      }
    }
  });
  return !length;
}
```
