---
id: 594810f028c0303b75339acf
title: Kitendakazi cha Ackermann
challengeType: 1
forumTopicId: 302223
dashedName: ackermann-function
---

# --description--

Kitendakazi cha Ackermann ni mfano wa kawaida wa kitendakazi kinachojirudia, kinachojulikana hasa kwa sababu si kitendakazi cha msingi kinachojirudia. Kinasonga kwa kasi sana kwa thamani, kama vile ukubwa wa mti wa wito wake.

Kitendakazi cha Ackermann kawaida hufafanuliwa kama ifuatavyo:

$A(m, n) = \\begin{cases} n+1 & \\mbox{if } m = 0 \\\\ A(m-1, 1) & \\mbox{if } m > 0 \\mbox{ and } n = 0 \\\\ A(m-1, A(m, n-1)) & \\mbox{if } m > 0 \\mbox{ and } n > 0. \\end{cases}$

Hoja zake hazijawahi kuwa hasi na daima humaliza.

# --instructions--

Andika kitendakazi kinachorudisha thamani ya $A(m, n)$. Usahihi wa hiari unapendekezwa (kwa sababu kitendakazi kinakua kwa kasi sana), lakini si sharti.

# --hints--

`ack` inapaswa kuwa kitendakazi.

```js
assert(typeof ack === 'function');
```

`ack(0, 0)` inapaswa kurudisha 1.

```js
assert(ack(0, 0) === 1);
```

`ack(1, 1)` inapaswa kurudisha 3.

```js
assert(ack(1, 1) === 3);
```

`ack(2, 5)` inapaswa kurudisha 13.

```js
assert(ack(2, 5) === 13);
```

`ack(3, 3)` inapaswa kurudisha 61.

```js
assert(ack(3, 3) === 61);
```

# --seed--

## --seed-contents--

```js
function ack(m, n) {

}
```

# --solutions--

```js
function ack(m, n) {
  return m === 0 ? n + 1 : ack(m - 1, n === 0 ? 1 : ack(m, n - 1));
}
```
