---
id: 594810f028c0303b75339ad5
title: Mchanganyiko wa Y
challengeType: 1
forumTopicId: 302345
dashedName: y-combinator
---

# --description--

Katika programu utendaji mkali <a href="https://www.freecodecamp.org/news/the-principles-of-functional-programming/" target="_blank" rel="noopener noreferrer nofollow">na hesabu ya lambda</a>, vitendakazi (usemi wa lambda) havina hali na vinaruhusiwa kurejelea tu hoja za vitendakazi vinavyozikwa. Hii inatofautisha na ufafanuzi wa kawaida wa kitendakazi cha kurudia ambapo kitendakazi kinaambatanishwa na hali ya kigezo na hali hii hutumika katika mwili wa kitendakazi.

Mchanganyiko wa Y ni kitendakazi isiyo na hali ambacho, kinapotumika kwa kitendakazi kingine isiyo na hali, hurudisha toleo la kurudia la kitendakazi hicho. Mchanganyiko wa Y ni rahisi zaidi kati ya darasa la vitendakazi kama hivi, vinavyoitwa mchanganyiko wa nukta thabiti.

# --instructions--

Fafanua kitendakazi isiyo na hali cha mchanganyiko wa Y na kitumie kuhesabu faktoria. Kitendakazi `factorial(N)` tayari kimekupa.

# --hints--

Y inapaswa kurudisha kitendakazi.

```js
assert.equal(typeof Y((f) => (n) => n), 'function');
```

factorial(1) inapaswa kurudisha 1.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(1), 1);
```

factorial(2) inapaswa kurudisha 2.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(2), 2);
```

factorial(3) inapaswa kurudisha 6.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(3), 6);
```

factorial(4) inapaswa kurudisha 24.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(4), 24);
```

factorial(10) inapaswa kurudisha 3628800.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(10), 3628800);
```

# --seed--

## --seed-contents--

```js
function Y(f) {
  return function() {

  };
}

var factorial = Y(function(f) {
  return function (n) {
    return n > 1 ? n * f(n - 1) : 1;
  };
});
```

# --solutions--

```js
var Y = f => (x => x(x))(y => f(x => y(y)(x)));
```
