---
id: 594810f028c0303b75339ad6
title: Wakilisho wa nambari za Zeckendorf
challengeType: 1
forumTopicId: 302346
dashedName: zeckendorf-number-representation
---

# --description--

Kama nambari zinaweza wakilishwa katika noti ya nafasi kama jumla ya mara nyingi ya nguvu za kumi (desimali) au mbili (binary); nambari zote kamili chanya zinaweza wakilishwa kama jumla ya mara moja au sifuri mara wanachama tofauti wa mfululizo wa Fibonacci. Kumbuka kwamba nambari sita za kwanza tofauti za Fibonacci ni: `1, 2, 3, 5, 8, 13`.

Nambari ya desimali kumi na moja inaweza kuandikwa kama `0*13 + 1*8 + 0*5 + 1*3 + 0*2 + 0*1` au `010100` katika noti ya nafasi ambapo safu zinawakilisha kuzidisha na mwanachama fulani wa mfululizo. Sifuri za mwanzoni huondolewa ili 11 desimali iwe `10100`. Hata hivyo 10100 siyo njia pekee ya kutengeneza 11 kutoka kwa nambari za Fibonacci, `0*13 + 1*8 + 0*5 + 0*3 + 1*2 + 1*1` au 010011 pia zingewakilisha 11 desimali. Kwa nambari halisi ya Zeckendorf kuna kikomo cha ziada kwamba *hakuna nambari mbili za mfululizo wa Fibonacci zinazoweza kutumika* jambo ambalo linasababisha suluhisho la kipekee la awali.

# --instructions--

Andika kitendakazi kinachozalisha na kurudisha wakilisho wa nambari za Zeckendorf za `n`.

# --hints--

`zeckendorf` inapaswa kuwa kitendakazi.

```js
assert.equal(typeof zeckendorf, 'function');
```

`zeckendorf(0)` inapaswa kurudisha `0`.

```js
assert.equal(zeckendorf(0), 0);

```

`zeckendorf(1)` inapaswa kurudisha `1`.

```js
assert.equal(zeckendorf(1), 1);
```

`zeckendorf(2)` inapaswa kurudisha `10`.

```js
assert.equal(zeckendorf(2), 10);
```

`zeckendorf(3)` inapaswa kurudisha `100`.

```js
assert.equal(zeckendorf(3), 100);
```

`zeckendorf(4)` inapaswa kurudisha `101`.

```js
assert.equal(zeckendorf(4), 101);
```

`zeckendorf(5)` inapaswa kurudisha `1000`.

```js
assert.equal(zeckendorf(5), 1000);
```

`zeckendorf(6)` inapaswa kurudisha `1001`.

```js
assert.equal(zeckendorf(6), 1001);
```

`zeckendorf(7)` inapaswa kurudisha `1010`.

```js
assert.equal(zeckendorf(7), 1010);
```

`zeckendorf(8)` inapaswa kurudisha `10000`.

```js
assert.equal(zeckendorf(8), 10000);
```

`zeckendorf(9)` inapaswa kurudisha `10001`.

```js
assert.equal(zeckendorf(9), 10001);
```

`zeckendorf(10)` inapaswa kurudisha `10010`.

```js
assert.equal(zeckendorf(10), 10010);
```

`zeckendorf(11)` inapaswa kurudisha `10100`.

```js
assert.equal(zeckendorf(11), 10100);
```

`zeckendorf(12)` inapaswa kurudisha `10101`.

```js
assert.equal(zeckendorf(12), 10101);
```

`zeckendorf(13)` inapaswa kurudisha `100000`.

```js
assert.equal(zeckendorf(13), 100000);
```

`zeckendorf(14)` inapaswa kurudisha `100001`.

```js
assert.equal(zeckendorf(14), 100001);
```

`zeckendorf(15)` inapaswa kurudisha `100010`.

```js
assert.equal(zeckendorf(15), 100010);
```

`zeckendorf(16)` inapaswa kurudisha `100100`.

```js
assert.equal(zeckendorf(16), 100100);
```

`zeckendorf(17)` inapaswa kurudisha `100101`.

```js
assert.equal(zeckendorf(17), 100101);
```

`zeckendorf(18)` inapaswa kurudisha `101000`.

```js
assert.equal(zeckendorf(18), 101000);
```

`zeckendorf(19)` inapaswa kurudisha `101001`.

```js
assert.equal(zeckendorf(19), 101001);
```

`zeckendorf(20)` inapaswa kurudisha `101010`.

```js
assert.equal(zeckendorf(20), 101010);
```

# --seed--

## --seed-contents--

```js
function zeckendorf(n) {

}
```

# --solutions--

```js
// zeckendorf :: Int -> Int
function zeckendorf(n) {
  const f = (m, x) => (m < x ? [m, 0] : [m - x, 1]);
  return parseInt((n === 0 ? ([0]) :
    mapAccumL(f, n, reverse(
      tail(fibUntil(n))
    ))[1]).join(''));
}

// fibUntil :: Int -> [Int]
let fibUntil = n => {
  const xs = [];
  until(
      ([a]) => a > n,
      ([a, b]) => (xs.push(a), [b, a + b]), [1, 1]
  );
  return xs;
};

let mapAccumL = (f, acc, xs) => (
  xs.reduce((a, x) => {
    const pair = f(a[0], x);

    return [pair[0], a[1].concat(pair[1])];
  }, [acc, []])
);

let until = (p, f, x) => {
  let v = x;
  while (!p(v)) v = f(v);
  return v;
};

const tail = xs => (
   xs.length ? xs.slice(1) : undefined
);

const reverse = xs => xs.slice(0).reverse();
```
