---
id: 594810f028c0303b75339ad2
title: Mzidishaji wa msalaba wa vekta
challengeType: 1
forumTopicId: 302342
dashedName: vector-cross-product
---

# --description--

Vekta hufafanuliwa kuwa na vipimo vitatu kwa kuwakilishwa na mkusanyiko uliopangwa wa nambari tatu: (X, Y, Z).

# --instructions--

Andika kitendakazi kinachopokea vekta mbili (safu za data) kama ingizo na kuhesabu mzidishaji wao wa msalaba. Kitendakazi chako kinapaswa kurudisha `null` kwa ingizo batili kama vile vekta zenye urefu tofauti.

# --hints--

`crossProduct` inapaswa kuwa kitendakazi.

```js
assert.equal(typeof crossProduct, 'function');
```

`crossProduct()` inapaswa kurudisha null.

```js
assert.equal(crossProduct(), null);
```

`crossProduct([1, 2, 3], [4, 5, 6])` inapaswa kurudisha `[-3, 6, -3]`.

```js
assert.deepEqual(res12, exp12);
```

# --before-each--

```js
const tv1 = [1, 2, 3];
const tv2 = [4, 5, 6];
const res12 = crossProduct(tv1, tv2);
const exp12 = [-3, 6, -3];
```

# --seed--

## --seed-contents--

```js
function crossProduct(a, b) {

}
```

# --solutions--

```js
function crossProduct(a, b) {
  if (!a || !b) {
    return null;
  }

  // Check lengths
  if (a.length !== 3 || b.length !== 3) {
    return null;
  }

  return [
    (a[1] * b[2]) - (a[2] * b[1]),
    (a[2] * b[0]) - (a[0] * b[2]),
    (a[0] * b[1]) - (a[1] * b[0])
  ];
}
```
