---
id: 594810f028c0303b75339ad1
title: Nambari za furaha
challengeType: 1
forumTopicId: 302280
dashedName: happy-numbers
---

# --description--

Nambari ya furaha hufafanuliwa kwa mchakato ufuatao:

Kuanzia na nambari yoyote kamili chanya, badilisha nambari hiyo kwa jumla ya mraba wa tarakimu zake, kisha rudia mchakato huu hadi nambari iwe sawa na `1` (ambayo itabaki hapo), au itazunguka bila mwisho katika mzunguko usiojumuisha `1`. Nambari ambazo mchakato huu unamalizika kwa `1` ni nambari za furaha, wakati zile ambazo hazimaliziki kwa `1` ni nambari zisizofuraha.

# --instructions--

Tekeleza kitendakazi kinachorejesha true ikiwa nambari ni ya furaha, au false ikiwa siyo.

# --hints--

`happy` inapaswa kuwa kitendakazi.

```js
assert(typeof happy === 'function');
```

`happy(1)` inapaswa kurejesha thamani ya kweli au siyo kweli.

```js
assert(typeof happy(1) === 'boolean');
```

`happy(1)` inapaswa kurejesha `true`.

```js
assert(happy(1));
```

`happy(2)` inapaswa kurejesha `false`.

```js
assert(!happy(2));
```

`happy(7)` inapaswa kurejesha `true`.

```js
assert(happy(7));
```

`happy(10)` inapaswa kurejesha `true`.

```js
assert(happy(10));
```

`happy(13)` inapaswa kurejesha `true`.

```js
assert(happy(13));
```

`happy(19)` inapaswa kurejesha `true`.

```js
assert(happy(19));
```

`happy(23)` inapaswa kurejesha `true`.

```js
assert(happy(23));
```

`happy(28)` inapaswa kurejesha `true`.

```js
assert(happy(28));
```

`happy(31)` inapaswa kurejesha `true`.

```js
assert(happy(31));
```

`happy(32)` inapaswa kurejesha `true`.

```js
assert(happy(32));
```

`happy(33)` inapaswa kurejesha `false`.

```js
assert(!happy(33));
```

# --seed--

## --seed-contents--

```js
function happy(number) {

}
```

# --solutions--

```js
function happy (number) {
  let m;
  let digit;
  const cycle = [];

  while (number !== 1 && cycle[number] !== true) {
    cycle[number] = true;
    m = 0;
    while (number > 0) {
      digit = number % 10;
      m += Math.pow(digit, 2);
      number = (number - digit) / 10;
    }
    number = m;
  }
  return (number === 1);
}
```
