---
id: 594810f028c0303b75339ad7
title: Algoriti ya kupunguza Zhang-Suen
challengeType: 1
forumTopicId: 302347
dashedName: zhang-suen-thinning-algorithm
---

# --description--

Hii ni algoriti inayotumika kupunguza picha za rangi nyeusi na nyeupe, yaani, biti moja kwa kila pikseli. Kwa mfano, kwa picha ya ingizo ya:

```js
const testImage1 = [
 '                               ',
 '#########       ########       ',
 '###   ####     ####  ####      ',
 '###    ###     ###    ###      ',
 '###   ####     ###             ',
 '#########      ###             ',
 '### ####       ###    ###      ',
 '###  ####  ### ####  #### ###  ',
 '###   #### ###  ########  ###  ',
 '                               '
];
```

Hutoa picha iliyopunguzwa kama ifuatavyo:

```js
[ '                               ',
  '########         ######        ',
  '#      #        ##             ',
  '#       #       #              ',
  '#      #        #              ',
  '###### #        #              ',
  '#     ##        #              ',
  '#      #    #   ##    ##   #   ',
  '#       #         ####         ',
  '                               ' ];
```

## Algoriti

Tafakari pikseli za rangi nyeusi ni moja na pikseli za rangi nyeupe ni sifuri, na picha ya ingizo ni safu ya mstatili ya N kwa M ya moja na sifuri. Algoriti hufanya kazi kwa pikseli zote nyeusi P1 ambazo zina majirani nane. Majirani hao, kwa mpangilio, yamepangwa kama ifuatavyo:

$$\begin{array}{|c|c|c|}
  \\hline
  P9 & P2              & P3\\\\ \\hline
  P8 & \boldsymbol{P1} & P4\\\\ \\hline
  P7 & P6              & P5\\\\ \\hline
\end{array}$$

Ni dhahiri kuwa pikseli za mipaka ya picha haziwezi kuwa na majirani nane kamili.

- Tafsiri $A(P1)$ = idadi ya mabadiliko kutoka nyeupe kwenda nyeusi, ($0 \to 1$) katika mfululizo P2, P3, P4, P5, P6, P7, P8, P9, P2. (Kumbuka P2 ya ziada mwishoni - ni mzunguko).
- Tafsiri $B(P1)$ = idadi ya majirani wa pikseli nyeusi wa P1. ($= \\sum(P2 \ldots P9)$)

**Hatua ya 1:**

Pikseli zote zinajaribiwa na zile zinazokidhi masharti yote yafuatayo (kwa wakati mmoja) zinatambuliwa tu katika hatua hii.

1. Pikseli ni nyeusi na ina majirani nane
2. $2 \le B(P1) \le 6$
3. $A(P1) = 1$
4. Angalau moja kati ya $P2$, $P4$ na $P6$ ni nyeupe
5. Angalau moja kati ya $P4$, $P6$ na $P8$ ni nyeupe

Baada ya kupitia picha na kukusanya pikseli zote zinazokidhi masharti yote ya hatua ya 1, pikseli zote zinazokidhi masharti haya zinabadilishwa kuwa nyeupe.

**Hatua ya 2:**

Pikseli zote zinajaribiwa tena na zile zinazokidhi masharti yote yafuatayo zinatambuliwa tu katika hatua hii.

1. Pikseli ni nyeusi na ina majirani nane
2. $2 \le B(P1) \le 6$
3. $A(P1) = 1$
4. Angalau moja kati ya $P2$, $P4$ na $P8$ ni nyeupe
5. Angalau moja kati ya $P2$, $P6$ na $P8$ ni nyeupe

Baada ya kupitia picha na kukusanya pikseli zote zinazokidhi masharti yote ya hatua ya 2, pikseli zote zinazokidhi masharti haya zinabadilishwa tena kuwa nyeupe.

**Kurudia:**

Kama kuna pikseli yoyote iliyobadilishwa katika mzunguko huu wa hatua ya 1 au hatua ya 2 basi hatua zote hurudiwa tena hadi hakuna pikseli yoyote ya picha kubadilishwa.

# --instructions--

Andika utaratibu wa kutekeleza kupunguza Zhang-Suen kwenye `image` uliotolewa, safu ya mfuatano wa herufi, ambapo kila mfuatano wa herufi unawakilisha mstari mmoja wa picha. Katika mfuatano wa herufi, `#` inawakilisha pikseli nyeusi, na nafasi tupu inawakilisha pikseli nyeupe. Kitendakazi kinapaswa kurudisha picha iliyopunguzwa, kwa kutumia uwakilishi uleule.

# --before-each--

```js
const _testImage1 = [
  '                               ',
  '#########       ########       ',
  '###   ####     ####  ####      ',
  '###    ###     ###    ###      ',
  '###   ####     ###             ',
  '#########      ###             ',
  '### ####       ###    ###      ',
  '###  ####  ### ####  #### ###  ',
  '###   #### ###  ########  ###  ',
  '                               '
];
const expected1 = [
  '                               ',
  '########         ######        ',
  '#      #        ##             ',
  '#       #       #              ',
  '#      #        #              ',
  '###### #        #              ',
  '#     ##        #              ',
  '#      #    #   ##    ##   #   ',
  '#       #         ####         ',
  '                               '
];
const _testImage2 = [
  '                                                          ',
  ' #################                   #############        ',
  ' ##################               ################        ',
  ' ###################            ##################        ',
  ' ########     #######          ###################        ',
  '   ######     #######         #######       ######        ',
  '   ######     #######        #######                      ',
  '   #################         #######                      ',
  '   ################          #######                      ',
  '   #################         #######                      ',
  '   ######     #######        #######                      ',
  '   ######     #######        #######                      ',
  '   ######     #######         #######       ######        ',
  ' ########     #######          ###################        ',
  ' ########     ####### ######    ################## ###### ',
  ' ########     ####### ######      ################ ###### ',
  ' ########     ####### ######         ############# ###### ',
  '                                                          '
];
const expected2 = [
  '                                                          ',
  '                                                          ',
  '    # ##########                       #######            ',
  '     ##        #                   ####       #           ',
  '     #          #                 ##                      ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     ############               #                         ',
  '     #          #               #                         ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     #                            ##                      ',
  '     #                             ############           ',
  '                       ###                          ###   ',
  '                                                          ',
  '                                                          '
];
```

# --hints--

`thinImage` inapaswa kuwa kitendakazi.

```js
assert.equal(typeof thinImage, 'function');
```

`thinImage` inapaswa kurudisha safu ya data.

```js
assert(Array.isArray(thinImage(_testImage1)));
```

`thinImage` inapaswa kurudisha safu ya mfuatano wa herufi.

```js
assert.equal(typeof thinImage(_testImage1)[0], 'string');
```

`thinImage(testImage1)` inapaswa kurudisha picha iliyopunguzwa kama katika mfano.

```js
assert.deepEqual(thinImage(_testImage1), expected1);
```

`thinImage(testImage2)` inapaswa kurudisha picha iliyopunguzwa.

```js
assert.deepEqual(thinImage(_testImage2), expected2);
```

# --seed--

## --seed-contents--

```js
function thinImage(image) {

}

const testImage1 = [
  '                               ',
  '#########       ########       ',
  '###   ####     ####  ####      ',
  '###    ###     ###    ###      ',
  '###   ####     ###             ',
  '#########      ###             ',
  '### ####       ###    ###      ',
  '###  ####  ### ####  #### ###  ',
  '###   #### ###  ########  ###  ',
  '                               '
];
```

# --solutions--

```js
function Point(x, y) {
  this.x = x;
  this.y = y;
}

const ZhangSuen = (function () {
  function ZhangSuen() {
  }

  ZhangSuen.nbrs = [[0, -1], [1, -1], [1, 0], [1, 1], [0, 1], [-1, 1], [-1, 0], [-1, -1], [0, -1]];

  ZhangSuen.nbrGroups = [[[0, 2, 4], [2, 4, 6]], [[0, 2, 6], [0, 4, 6]]];

  ZhangSuen.toWhite = [];

  ZhangSuen.main = function (image) {
    ZhangSuen.grid = new Array(image);
    for (let r = 0; r < image.length; r++) {
      ZhangSuen.grid[r] = image[r].split('');
    }
    ZhangSuen.thinImage();
    return ZhangSuen.getResult();
  };

  ZhangSuen.thinImage = function () {
    let firstStep = false;
    let hasChanged;
    do {
      hasChanged = false;
      firstStep = !firstStep;
      for (let r = 1; r < ZhangSuen.grid.length - 1; r++) {
        for (let c = 1; c < ZhangSuen.grid[0].length - 1; c++) {
          if (ZhangSuen.grid[r][c] !== '#') {
            continue;
          }
          const nn = ZhangSuen.numNeighbors(r, c);
          if (nn < 2 || nn > 6) {
            continue;
          }
          if (ZhangSuen.numTransitions(r, c) !== 1) {
            continue;
          }
          if (!ZhangSuen.atLeastOneIsWhite(r, c, firstStep ? 0 : 1)) {
            continue;
          }
          ZhangSuen.toWhite.push(new Point(c, r));
          hasChanged = true;
        }
      }
      for (let i = 0; i < ZhangSuen.toWhite.length; i++) {
        const p = ZhangSuen.toWhite[i];
        ZhangSuen.grid[p.y][p.x] = ' ';
      }
      ZhangSuen.toWhite = [];
    } while ((firstStep || hasChanged));
  };

  ZhangSuen.numNeighbors = function (r, c) {
    let count = 0;
    for (let i = 0; i < ZhangSuen.nbrs.length - 1; i++) {
      if (ZhangSuen.grid[r + ZhangSuen.nbrs[i][1]][c + ZhangSuen.nbrs[i][0]] === '#') {
        count++;
      }
    }
    return count;
  };

  ZhangSuen.numTransitions = function (r, c) {
    let count = 0;
    for (let i = 0; i < ZhangSuen.nbrs.length - 1; i++) {
      if (ZhangSuen.grid[r + ZhangSuen.nbrs[i][1]][c + ZhangSuen.nbrs[i][0]] === ' ') {
        if (ZhangSuen.grid[r + ZhangSuen.nbrs[i + 1][1]][c + ZhangSuen.nbrs[i + 1][0]] === '#') {
          count++;
        }
      }
    }
    return count;
  };

  ZhangSuen.atLeastOneIsWhite = function (r, c, step) {
    let count = 0;
    const group = ZhangSuen.nbrGroups[step];
    for (let i = 0; i < 2; i++) {
      for (let j = 0; j < group[i].length; j++) {
        const nbr = ZhangSuen.nbrs[group[i][j]];
        if (ZhangSuen.grid[r + nbr[1]][c + nbr[0]] === ' ') {
          count++;
          break;
        }
      }
    }
    return count > 1;
  };

  ZhangSuen.getResult = function () {
    const result = [];
    for (let i = 0; i < ZhangSuen.grid.length; i++) {
      const row = ZhangSuen.grid[i].join('');
      result.push(row);
    }
    return result;
  };
  return ZhangSuen;
}());

function thinImage(image) {
  return ZhangSuen.main(image);
}
```
