---
id: 594810f028c0303b75339acb
title: Milango 100
challengeType: 1
forumTopicId: 302217
dashedName: 100-doors
---

# --description--

Kuna milango 100 mfululizo ambayo yote kwa awali imefungwa. Unapita milango hiyo mara 100. Mara ya kwanza unapita kila mlango na 'badilisha' hali ya mlango (ikiwa mlango umefungwa, ufungue; ikiwa umefunguliwa, ufunge). Mara ya pili, tembelea kila mlango wa nambari 2 tu (yaani, mlango nambari 2, 4, 6, ...) na ubadilishe hali yake. Mara ya tatu, tembelea kila mlango wa nambari 3 (yaani, mlango nambari 3, 6, 9, ...), n.k., hadi utakapotembelea mlango wa nambari 100 pekee.

# --instructions--

Tekeleza kitendakazi ili kubaini hali ya milango baada ya kupita mara ya mwisho. Rudisha matokeo ya mwisho katika safu ya data, ukiweka nambari ya mlango tu ikiwa mlango huo umefunguliwa.

# --hints--

`getFinalOpenedDoors` inapaswa kuwa kitendakazi.

```js
assert(typeof getFinalOpenedDoors === 'function');
```

`getFinalOpenedDoors` inapaswa kurudisha safu ya data.

```js
assert(Array.isArray(getFinalOpenedDoors(100)));
```

`getFinalOpenedDoors` inapaswa kutoa matokeo sahihi.

```js
const solution = [1, 4, 9, 16, 25, 36, 49, 64, 81, 100];
assert.deepEqual(getFinalOpenedDoors(100), solution);
```

# --seed--

## --seed-contents--

```js
function getFinalOpenedDoors(numDoors) {

}
```

# --solutions--

```js
function getFinalOpenedDoors(numDoors) {
  // this is the final pattern (always squares).
  // thus, the most efficient solution simply returns an array of squares up to numDoors).
  const finalState = [];
  let i = 1;
  while (Math.pow(i, 2) <= numDoors) {
    finalState.push(Math.pow(i, 2));
    i++;
  }
  return finalState;
}
```
