---
id: 5949b579404977fbaefcd736
title: nambari bilioni 9 za Mungu nambari kamili
challengeType: 1
forumTopicId: 302219
dashedName: 9-billion-names-of-god-the-integer
---

# --description--

Zoezi hili ni toleo la hadithi fupi ya Arthur C. Clarke.

(Watatekelezaji wanapaswa kuwa na ufahamu wa matokeo ya kukamilisha zoezi hili.)

Kwa undani, kufafanua maana ya "jina":

<ul>
  <li>Nambari kamili 1 ina jina 1 "1".</li>
  <li>Nambari kamili 2 ina majina 2 "1+1" na "2".</li>
  <li>Nambari kamili 3 ina majina 3 "1+1+1", "2+1", na "3".</li>
  <li>Nambari kamili 4 ina majina 5 "1+1+1+1", "2+1+1", "2+2", "3+1", "4".</li>
  <li>Nambari kamili 5 ina majina 7 "1+1+1+1+1", "2+1+1+1", "2+2+1", "3+1+1", "3+2", "4+1", "5".</li>
</ul>

Hii inaweza kuonyeshwa kwa muundo ufuatao:

<pre>          1
        1   1
      1   1   1
    1   2   1   1
  1   2   2   1   1
1   3   3   2   1   1
</pre>

Ambapo safu $n$ inahusiana na nambari kamili $n$, na kila safu $m$ nguzo $C$ kutoka kushoto kwenda kulia inahusiana na idadi ya majina yanayoanza na $C$.

Kwa hiari tambua kuwa jumla ya safu ya $n$-th $P(n)$ ni kitendakazi cha mgawanyo wa nambari kamili.

# --instructions--

Tekeleza kitendakazi kinachorudisha jumla ya safu ya $n$-th.

# --hints--

`numberOfNames` inapaswa kuwa kitendakazi.

```js
assert(typeof numberOfNames === 'function');
```

`numberOfNames(5)` inapaswa kuwa sawa na 7.

```js
assert.equal(numberOfNames(5), 7);
```

`numberOfNames(12)` inapaswa kuwa sawa na 77.

```js
assert.equal(numberOfNames(12), 77);
```

`numberOfNames(18)` inapaswa kuwa sawa na 385.

```js
assert.equal(numberOfNames(18), 385);
```

`numberOfNames(23)` inapaswa kuwa sawa na 1255.

```js
assert.equal(numberOfNames(23), 1255);
```

`numberOfNames(42)` inapaswa kuwa sawa na 53174.

```js
assert.equal(numberOfNames(42), 53174);
```

`numberOfNames(123)` inapaswa kuwa sawa na 2552338241.

```js
assert.equal(numberOfNames(123), 2552338241);
```

# --seed--

## --seed-contents--

```js
function numberOfNames(num) {

  return true;
}
```

# --solutions--

```js
function numberOfNames(num) {
  const cache = [
    [1]
  ];
  for (let l = cache.length; l < num + 1; l++) {
    let Aa;
    let Mi;
    const r = [0];
    for (let x = 1; x < l + 1; x++) {
      r.push(r[r.length - 1] + (Aa = cache[l - x < 0 ? cache.length - (l - x) : l - x])[(Mi = Math.min(x, l - x)) < 0 ? Aa.length - Mi : Mi]);
    }
    cache.push(r);
  }
  return cache[num][cache[num].length - 1];
}
```
