---
id: 594d8d0ab97724821379b1e6
title: Wastani/Mode
challengeType: 1
forumTopicId: 302226
dashedName: averagesmode
---

# --description--

Andika kitendakazi `mode` ili kupata thamani inayojitokeza mara nyingi zaidi katika safu ya data.

Hali ambapo mkusanyiko ni tupu inaweza kupuuzwa. Tahadhari inapaswa kuchukuliwa kushughulikia hali ambapo mode si ya kipekee.

Ikiwa si sahihi au haiwezekani kuunga mkono mkusanyiko wa jumla, tumia vector (safu ya data), ikiwa inawezekana. Ikiwa si sahihi au haiwezekani kuunga mkono aina isiyoelezwa ya thamani, tumia nambari kamili.

# --before-each--

```js
const arr1 = [1, 3, 6, 6, 6, 6, 7, 7, 12, 12, 17];
const arr2 = [1, 2, 4, 4, 1];
```

# --hints--

`mode` inapaswa kuwa kitendakazi.

```js
assert(typeof mode === 'function');
```

`mode([1, 3, 6, 6, 6, 6, 7, 7, 12, 12, 17])` inapaswa kuwa sawa na `[6]`

```js
assert.deepEqual(mode(arr1), [6]);
```

`mode([1, 2, 4, 4, 1])` inapaswa kuwa sawa na `[1, 4]`.

```js
assert.deepEqual(mode(arr2).sort(), [1, 4]);
```

# --seed--

## --seed-contents--

```js
function mode(arr) {

  return true;
}
```

# --solutions--

```js
function mode(arr) {
  const counter = {};
  let result = [];
  let max = 0;
  // for (const i in arr) {
  arr.forEach(el => {
    if (!(el in counter)) {
      counter[el] = 0;
    }
    counter[el]++;

    if (counter[el] === max) {
      result.push(el);
    }
    else if (counter[el] > max) {
      max = counter[el];
      result = [el];
    }
  });
  return result;
}
```
