---
id: 594da033de4190850b893874
title: Wastani/mzizi wa mraba wa wastani
challengeType: 1
forumTopicId: 302228
dashedName: averagesroot-mean-square
---

# --description--

Hesabu mzizi wa mraba wa wastani (RMS) wa nambari kutoka 1 hadi 10 pamoja.

<abbr title="Root mean square">RMS</abbr> huhesabiwa kwa kuchukua mzizi wa mraba wa wastani wa mraba za nambari, kama inavyoonyeshwa na mlinganyo:

$$x\_{\\mathrm{rms}} = \\sqrt {{{x_1}^2 + {x_2}^2 + \\cdots + {x_n}^2} \\over n}. $$

# --hints--

`rms` inapaswa kuwa kitendakazi.

```js
assert(typeof rms === 'function');
```

`rms([1, 2, 3, 4, 5, 6, 7, 8, 9, 10])` inapaswa kuwa sawa na `6.2048368229954285`.

```js
const arr1 = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const answer1 = 6.2048368229954285;
assert.equal(rms(arr1), answer1);
```

# --seed--

## --seed-contents--

```js
function rms(arr) {

}
```

# --solutions--

```js
function rms(arr) {
  const sumOfSquares = arr.reduce((s, x) => s + x * x, 0);
  return Math.sqrt(sumOfSquares / arr.length);
}
```
