---
id: 594dc6c729e5700999302b45
title: Mabano yaliyo sawa
challengeType: 1
forumTopicId: 302230
dashedName: balanced-brackets
---

# --description--

Tambua kama mfuatano wa mabano uliotengenezwa uko sawa; yaani, kama unajumuisha jozi zote za mabano ya kufungua/kufunga (kwa mpangilio huo), ambayo hakuna hata moja iliyopangwa vibaya ndani yake.

**Mifano:**
| Ingizo                    | Matokeo |
| ------------------------- | -------- |
| <code>\[]</code>          | true     |
| <code>]\[</code>          | false    |
| <code>[][]</code>         | true     |
| <code>]\[]</code>         | false    |
| <code>\[]]\[\[]</code>    | false    |
| <code>\[\[\[\[]]]]</code> | true     |

# --before-each--

```js
const testCases = [
  '[]',
  ']][[[][][][]][',
  '[][[[[][][[[]]]]]]',
  '][',
  '[[[]]]][[]',
  '][[]',
  '][[][]][[[]]',
  '[[][]]][',
  '[[[]]][[]]]][][[',
  '[]][[]]][[[[][]]',
  '][]][[][',
  '[[]][[][]]',
  '[[]]',
  ']][]][[]][[[',
  '][]][][[',
  '][][',
  '[]]]',
  ''
];
```

# --hints--

`isBalanced` inapaswa kuwa kitendakazi.

```js
assert(typeof isBalanced === 'function');
```

`isBalanced("[]")` inapaswa kurudisha true.

```js
assert(isBalanced(testCases[0]));
```

`isBalanced("]][[[][][][]][")` inapaswa kurudisha false.

```js
assert(!isBalanced(testCases[1]));
```

`isBalanced("[][[[[][][[[]]]]]]")` inapaswa kurudisha true.

```js
assert(isBalanced(testCases[2]));
```

`isBalanced("][")` inapaswa kurudisha false.

```js
assert(!isBalanced(testCases[3]));
```

`isBalanced("[[[]]]][[]")` inapaswa kurudisha false.

```js
assert(!isBalanced(testCases[4]));
```

`isBalanced("][[]")` inapaswa kurudisha false.

```js
assert(!isBalanced(testCases[5]));
```

`isBalanced("][[][]][[[]]")` inapaswa kurudisha false.

```js
assert(!isBalanced(testCases[6]));
```

`isBalanced("[[][]]][")` inapaswa kurudisha false.

```js
assert(!isBalanced(testCases[7]));
```

`isBalanced("[[[]]][[]]]][][[")` inapaswa kurudisha false.

```js
assert(!isBalanced(testCases[8]));
```

`isBalanced("[]][[]]][[[[][]]")` inapaswa kurudisha false.

```js
assert(!isBalanced(testCases[9]));
```

`isBalanced("][]][[][")` inapaswa kurudisha false.

```js
assert(!isBalanced(testCases[10]));
```

`isBalanced("[[]][[][]]")` inapaswa kurudisha true.

```js
assert(isBalanced(testCases[11]));
```

`isBalanced("[[]]")` inapaswa kurudisha true.

```js
assert(isBalanced(testCases[12]));
```

`isBalanced("]][]][[]][[[")` inapaswa kurudisha false.

```js
assert(!isBalanced(testCases[13]));
```

`isBalanced("][]][][[")` inapaswa kurudisha false.

```js
assert(!isBalanced(testCases[14]));
```

`isBalanced("][][")` inapaswa kurudisha false.

```js
assert(!isBalanced(testCases[15]));
```

`isBalanced("[]]]")` inapaswa kurudisha false.

```js
assert(!isBalanced(testCases[16]));
```

`isBalanced("")` inapaswa kurudisha true.

```js
assert(isBalanced(testCases[17]));
```

# --seed--

## --seed-contents--

```js
function isBalanced(str) {

  return true;
}
```

# --solutions--

```js
function isBalanced(str) {
  if (str === '') return true;
  let a = str;
  let b;
  do {
    b = a;
    a = a.replace(/\[\]/g, '');
  } while (a !== b);
  return !a;
}
```
