---
id: 594ecc0d9a8cf816e3340187
title: Nambari za taxicab
challengeType: 1
forumTopicId: 302337
dashedName: taxicab-numbers
---

# --description--

Nambari ya taxicab (tafsiri inayotumika hapa) ni nambari kamili chanya inayoweza kudhihirishwa kama jumla ya makubaliano mawili ya nambari kamili chanya kwa njia zaidi ya moja.

Nambari ya kwanza ya taxicab ni `1729`, ambayo ni:

1<sup>3</sup> + 12<sup>3</sup> na

9<sup>3</sup> + 10<sup>3</sup>.

Nambari za taxicab pia zinajulikana kama:

<ul>
  <li>nambari za taxi</li>
  <li>nambari za taxi-cab</li>
  <li>nambari za taxi cab</li>
  <li>nambari za Hardy-Ramanujan</li>
</ul>

# --instructions--

Andika kitendakazi kinachorudisha nambari za chini kabisa za taxicab `n`. Kwa kila nambari ya taxicab, onyesha nambari hiyo pamoja na makubaliano yake ya makubaliano.

# --before-each--

```js
const res4 = [1729, 4104, 13832, 20683];
const res25 = [
  1729, 4104, 13832, 20683, 32832, 39312, 40033, 46683, 64232, 65728, 110656,
  110808, 134379, 149389, 165464, 171288, 195841, 216027, 216125, 262656, 314496, 320264, 327763,
  373464, 402597
];

const res39From20To29 = [314496, 320264, 327763, 373464, 402597, 439101, 443889, 513000, 513856];
```

# --hints--

`taxicabNumbers` inapaswa kuwa kitendakazi.

```js
assert(typeof taxicabNumbers === 'function');
```

`taxicabNumbers` inapaswa kurudisha safu ya data.

```js
assert(typeof taxicabNumbers(2) === 'object');
```

`taxicabNumbers` inapaswa kurudisha safu ya nambari.

```js
assert(typeof taxicabNumbers(100)[0] === 'number');
```

`taxicabNumbers(4)` inapaswa kurudisha [1729, 4104, 13832, 20683].

```js
assert.deepEqual(taxicabNumbers(4), res4);
```

`taxicabNumbers(25)` inapaswa kurudisha [1729, 4104, 13832, 20683, 32832, 39312, 40033, 46683, 64232, 65728, 110656, 110808, 134379, 149389, 165464, 171288, 195841, 216027, 216125, 262656, 314496, 320264, 327763, 373464, 402597]

```js
assert.deepEqual(taxicabNumbers(25), res25);
```

`taxicabNumbers(39)` nambari zinazotokana kutoka 20 - 29 zinapaswa kuwa [314496,320264,327763,373464,402597,439101,443889,513000,513856].

```js
assert.deepEqual(taxicabNumbers(39).slice(20, 29), res39From20To29);
```

# --seed--

## --seed-contents--

```js
function taxicabNumbers(n) {

  return true;
}
```

# --solutions--

```js
function taxicabNumbers(nNumbers) {
  const cubeN = [];
  const s3s = {};

  const e = 100;
  for (let n = 1; n < e; n += 1) {
    cubeN[n] = n * n * n;
  }

  for (let a = 1; a < e - 1; a += 1) {
    const a3 = cubeN[a];
    for (let b = a; b < e; b += 1) {
      const b3 = cubeN[b];
      const s3 = a3 + b3;

      let abs = s3s[s3];
      if (!abs) {
        s3s[s3] = abs = [];
      }
      abs.push([a, b]);
    }
  }

  let i = 0;
  const res = [];
  Object.keys(s3s).forEach(s3 => {
    const abs = s3s[s3];
    if (abs.length >= 2) { // No two cube pairs found
      i += 1;
      if (i <= nNumbers) {
        res.push(s3);
      }
    }
  });
  return res.map(item => parseInt(item, 10));
}
```
