---
id: 594fa2746886f41f7d8bf225
title: kupanga kwa topolojia
challengeType: 1
forumTopicId: 302340
dashedName: topological-sort
---

# --description--

Ikiwa umepewa ramani kati ya vitu, na vitu vinavyotegemea, kupanga kwa topolojia huweka vitu kwa mpangilio usiopelekea kitu chochote kuzingatia kitu ambacho kinategemea. Kuna algoriti maarufu mbili za kupanga kwa topolojia: kupanga kwa topolojia ya Kahn (1962) na utafutaji wa kina kwanza.

# --instructions--

Andika kitendakazi kitakachorudisha orodha yenye mpangilio halali wa kukusanya maktaba kutoka kwa utegemezi wake.

- Kubali majina ya maktaba kuwa maneno moja.
- Vitu vilivyotajwa kama vinategemea tu havina utegemezi wao, lakini mpangilio wao wa kukusanya lazima utolewe.
- Utegemezi wowote wa kujitegemea usizingatiwe.
- Utegemezi usio na mpangilio wowote usizingatiwe.

Tumia data ifuatayo kama mfano:

<pre>
MAKTABA          UTEGEMEZI WA MAKTABA
=======          ====================
des_system_lib   std synopsys std_cell_lib des_system_lib dw02 dw01 ramlib ieee
dw01             ieee dw01 dware gtech
dw02             ieee dw02 dware
dw03             std synopsys dware dw03 dw02 dw01 ieee gtech
dw04             dw04 ieee dw01 dware gtech
dw05             dw05 ieee dware
dw06             dw06 ieee dware
dw07             ieee dware
dware            ieee dware
gtech            ieee gtech
ramlib           std ieee
std_cell_lib     ieee std_cell_lib
synopsys
</pre>

Kukusanya maktaba katika lugha ya VHDL kuna sharti kwamba maktaba lazima ikusanywe baada ya maktaba yoyote inayotegemea. Data iliyo juu isingeweza kupangwa ikiwa, kwa mfano, `dw04` itaongezwa kwenye orodha ya utegemezi wa `dw01`.

Ingizo la kitendakazi litakuwa mfuatano wa herufi wenye mistari mingi, kila mstari utakuwa na jina la maktaba, ikifuatiwa na utegemezi wake (ikiwa upo).

Kwa mfano:

```js
const libsSimple =
  `aaa bbb
  bbb`;
```

# --before-each--

```js
const libsSimple =
  `aaa bbb
  bbb`;

const libsVHDL =
  `des_system_lib   std synopsys std_cell_lib des_system_lib dw02 dw01 ramlib ieee
  dw01             ieee dw01 dware gtech
  dw02             ieee dw02 dware
  dw03             std synopsys dware dw03 dw02 dw01 ieee gtech
  dw04             dw04 ieee dw01 dware gtech
  dw05             dw05 ieee dware
  dw06             dw06 ieee dware
  dw07             ieee dware
  dware            ieee dware
  gtech            ieee gtech
  ramlib           std ieee
  std_cell_lib     ieee std_cell_lib
  synopsys`;

const libsCustom =
  `a b c d
  b c d
  d c
  c base
  base`;

const libsUnorderable =
  `TestLib Base MainLib
  MainLib TestLib
  Base`;
```

# --hints--

`topologicalSort` inapaswa kuwa kitendakazi.

```js
assert(typeof topologicalSort === 'function');
```

`topologicalSort(libsSimple)` inapaswa kurudisha safu ya data.

```js
assert(Array.isArray(topologicalSort(libsSimple)));
```

`topologicalSort(libsSimple)` inapaswa kurudisha `['bbb', 'aaa']`.

```js
assert.deepEqual(topologicalSort(libsSimple), ['bbb', 'aaa']);
```

`topologicalSort(libsVHDL)` inapaswa kurudisha `['ieee', 'std_cell_lib', 'gtech', 'dware', 'dw07', 'dw06', 'dw05', 'dw02', 'dw01', 'dw04', 'std', 'ramlib', 'synopsys', 'dw03', 'des_system_lib']`.

```js
assert.deepEqual(topologicalSort(libsVHDL), ['ieee', 'std_cell_lib', 'gtech', 'dware', 'dw07', 'dw06', 'dw05', 'dw02', 'dw01', 'dw04', 'std', 'ramlib', 'synopsys', 'dw03', 'des_system_lib']);
```

`topologicalSort(libsCustom)` inapaswa kurudisha `['base', 'c', 'd', 'b', 'a']`.

```js
assert.deepEqual(topologicalSort(libsCustom), ['base', 'c', 'd', 'b', 'a']);
```

`topologicalSort` inapaswa kupuuza utegemezi usio na mpangilio.

```js
assert.deepEqual(topologicalSort(libsUnorderable), ['Base']);
```

# --seed--

## --seed-contents--

```js
function topologicalSort(libs) {

  return true;
}
```

# --solutions--

```js
function topologicalSort(libs) {
  // A map of the input data, with the keys as the packages, and the values as
  // and array of packages on which it depends.
  const D = libs
    .split('\n')
    .map(e => e.split(' ').filter(ep => ep !== ''))
    .reduce((p, c) =>
      p.set(c[0], c.filter((e, i) => (i > 0 && e !== c[0] ? e : null))), new Map());
  [].concat(...D.values()).forEach(e => {
    D.set(e, D.get(e) || []);
  });

  // The above map rotated so that it represents a DAG of the form
  // Map {
  //    A => [ A, B, C],
  //    B => [C],
  //    C => []
  // }
  // where each key represents a node, and the array contains the edges.
  const G = [...D.keys()].reduce((p, c) =>
    p.set(
      c,
      [...D.keys()].filter(e => D.get(e).includes(c))),
    new Map()
  );

  // An array of leaf nodes; nodes with 0 in degrees.
  const Q = [...D.keys()].filter(e => D.get(e).length === 0);

  // The result array.
  const S = [];
  while (Q.length) {
    const u = Q.pop();
    S.push(u);
    G.get(u).forEach(v => {
      D.set(v, D.get(v).filter(e => e !== u));
      if (D.get(v).length === 0) {
        Q.push(v);
      }
    });
  }

  return S;
}
```
