---
id: 5951815dd895584b06884620
title: Mizunguko yenye mduara uliotolewa kupitia pointi mbili
challengeType: 1
forumTopicId: 302231
dashedName: circles-of-given-radius-through-two-points
---

# --description--

Ikiwa na pointi mbili kwenye ndege na mduara, kawaida mizunguko miwili yenye mduara uliotolewa inaweza kuchorwa kupitia pointi hizo.

**Mambo ya kuzingatia:**

<ul>
  <li>Mduara wa sifuri unapaswa kutambuliwa kama hauelezi mizunguko (isipokuwa katika hali ambapo pointi hizo ni sawa).</li>
  <li>Kama pointi hizo ni sawa basi kuna mizunguko isiyo na kikomo yenye pointi hiyo kwenye mduara wake inaweza kuchorwa, isipokuwa mduara pia ni sifuri ambayo basi huleta mizunguko kuwa pointi moja.</li>
  <li>Kama pointi hizo zinaunda kipenyo basi rudisha mduara mmoja tu.</li>
  <li>Kama pointi hizo ziko mbali sana basi hakuna mizunguko inaweza kuchorwa.</li>
</ul>

# --instructions--

Tekeleza kitendakazi kinachopokea pointi mbili na mduara na kurudisha mizunguko miwili kupitia pointi hizo. Kwa kila mduara utakaopatikana, toa kuratibu za katikati ya kila mduara zenye kuzungushwa hadi tarakimu nne za desimali. Rudisha kila kuratibu kama safu ya data, na kuratibu kama safu ya safu za data.

**Kwa hali za kipekee, rudisha yafuatayo:**

<ul>
  <li>Kama pointi ziko kwenye kipenyo, rudisha pointi moja. Hata hivyo, kama mduara pia ni sifuri, rudisha <code>"Radius Zero"</code>.</li>
  <li>Kama pointi ni sawa, rudisha <code>"Coincident point. Infinite solutions"</code>.</li>
  <li>Kama pointi ziko mbali zaidi ya kipenyo, rudisha <code>"No intersection. Points further apart than circle diameter"</code>.</li>
</ul>

**Mifano ya ingizo:**

<pre>      p1                p2           r
0.1234, 0.9876    0.8765, 0.2345    2.0
0.0000, 2.0000    0.0000, 0.0000    1.0
0.1234, 0.9876    0.1234, 0.9876    2.0
0.1234, 0.9876    0.8765, 0.2345    0.5
0.1234, 0.9876    0.1234, 0.9876    0.0
</pre>

# --before-each--

```js
const testCases = [
  [[0.1234, 0.9876], [0.8765, 0.2345], 2.0],
  [[0.0000, 2.0000], [0.0000, 0.0000], 1.0],
  [[0.1234, 0.9876], [0.1234, 0.9876], 2.0],
  [[0.1234, 0.9876], [0.8765, 0.2345], 0.5],
  [[0.1234, 0.9876], [0.1234, 0.9876], 0.0]
];
const answers = [
  [[1.8631, 1.9742], [-0.8632, -0.7521]],
  [0, 1],
  'Coincident point. Infinite solutions',
  'No intersection. Points further apart than circle diameter',
  'Radius Zero'
];
```

# --hints--

`getCircles` inapaswa kuwa kitendakazi.

```js
assert(typeof getCircles === 'function');
```

`getCircles([0.1234, 0.9876], [0.8765, 0.2345], 2.0)` inapaswa kurudisha `[[1.8631, 1.9742], [-0.8632, -0.7521]]`.

```js
assert.deepEqual(getCircles(...testCases[0]), answers[0]);
```

`getCircles([0.0000, 2.0000], [0.0000, 0.0000], 1.0)` inapaswa kurudisha `[0, 1]`

```js
assert.deepEqual(getCircles(...testCases[1]), answers[1]);
```

`getCircles([0.1234, 0.9876], [0.1234, 0.9876], 2.0)` inapaswa kurudisha `Coincident point. Infinite solutions`

```js
assert.deepEqual(getCircles(...testCases[2]), answers[2]);
```

`getCircles([0.1234, 0.9876], [0.8765, 0.2345], 0.5)` inapaswa kurudisha `No intersection. Points further apart than circle diameter`

```js
assert.deepEqual(getCircles(...testCases[3]), answers[3]);
```

`getCircles([0.1234, 0.9876], [0.1234, 0.9876], 0.0)` inapaswa kurudisha `Radius Zero`

```js
assert.deepEqual(getCircles(...testCases[4]), answers[4]);
```

# --seed--

## --seed-contents--

```js
function getCircles(...args) {

  return true;
}
```

# --solutions--

```js
const hDist = (p1, p2) => Math.hypot(...p1.map((e, i) => e - p2[i])) / 2;
const pAng = (p1, p2) => Math.atan(p1.map((e, i) => e - p2[i]).reduce((p, c) => c / p, 1));
const solveF = (p, r) => t => [parseFloat((r * Math.cos(t) + p[0]).toFixed(4)), parseFloat((r * Math.sin(t) + p[1]).toFixed(4))];
const diamPoints = (p1, p2) => p1.map((e, i) => parseFloat((e + (p2[i] - e) / 2).toFixed(4)));

function getCircles(...args) {
  const [p1, p2, s] = args;
  const solve = solveF(p1, s);
  const halfDist = hDist(p1, p2);

  let msg = [];
  switch (Math.sign(s - halfDist)) {
    case 0:
      msg = s ? diamPoints(p1, p2) :
        'Radius Zero';
      break;
    case 1:
      if (!halfDist) {
        msg = 'Coincident point. Infinite solutions';
      }
      else {
        const theta = pAng(p1, p2);
        const theta2 = Math.acos(halfDist / s);
        [1, -1].map(e => solve(theta + e * theta2)).forEach(
          e => msg.push(e));
      }
      break;
    case -1:
      msg = 'No intersection. Points further apart than circle diameter';
      break;
    default:
      msg = 'Reached the default';
  }
  return msg;
}
```
