---
id: 595668ca4cfe1af2fb9818d4
title: Mfululizo wa Harshad au Niven
challengeType: 1
forumTopicId: 302281
dashedName: harshad-or-niven-series
---

# --description--

Nambari za Harshad au Niven ni nambari kamili chanya ≥ 1 ambazo zinaweza kugawika kwa jumla ya tarakimu zao bila mabaki.

Kwa mfano, `42` ni nambari ya Harshad kwa sababu `42` inagawika kwa `(4 + 2)` bila mabaki.

Kubali kwamba mfululizo umefafanuliwa kama nambari katika mpangilio wa kuongezeka.

# --instructions--

Tekeleza kitendakazi cha kuzalisha wanachama wa mfululizo wa Harshad mfululizo.

Kitumie kurudisha safu ya data yenye wanachama kumi wa mfululizo, kuanzia na nambari ya kwanza ya Harshad inayozidi `n`.

# --hints--

`isHarshadOrNiven` inapaswa kuwa kitendakazi.

```js
assert(typeof isHarshadOrNiven === 'function');
```

`isHarshadOrNiven(10)` inapaswa kurudisha `[12, 18, 20, 21, 24, 27, 30, 36, 40, 42]`

```js
assert.deepEqual(isHarshadOrNiven(10), [12, 18, 20, 21, 24, 27, 30, 36, 40, 42]);
```

`isHarshadOrNiven(400)` inapaswa kurudisha `[402, 405, 407, 408, 410, 414, 420, 423, 432, 440]`

```js
assert.deepEqual(isHarshadOrNiven(400), [402, 405, 407, 408, 410, 414, 420, 423, 432, 440]);
```

`isHarshadOrNiven(1000)` inapaswa kurudisha `[1002, 1008, 1010, 1011, 1012, 1014, 1015, 1016, 1017, 1020]`

```js
assert.deepEqual(isHarshadOrNiven(1000), [1002, 1008, 1010, 1011, 1012, 1014, 1015, 1016, 1017, 1020]);
```

# --seed--

## --seed-contents--

```js
function isHarshadOrNiven(n) {
  const res = [];

  return res;
}
```

# --solutions--

```js
function isHarshadOrNiven(n) {
  function isHarshad(num) {
    let s = 0;
    const nStr = num.toString();
    for (let i = 0; i < nStr.length; ++i) {
      s += parseInt(nStr.charAt(i), 10);
    }
    return n % s === 0;
  }

  const res = [];
  let count = 0;

  while (count < 10) {
    n++;
    if (isHarshad(n)) {
      count++;
      res.push(n);
    }
  }

  return res;
}
```
