---
id: 59622f89e4e137560018a40e
title: Mfululizo wa Hofstadter Figure-Figure
challengeType: 1
forumTopicId: 302286
dashedName: hofstadter-figure-figure-sequences
---

# --description--

Mfululizo wa Hofstadter Figure-Figure $R_n$ na $S_n$ unatolewa na

$R_1 = 1\\ ;\\ S_1 = 2 \\\\R_n = R_{n-1} + S_{n-1}, \\quad n>1.$

Hasa, mfululizo $R_n$ una thamani

<pre>1, 3, 7, 12, 18, ...</pre>

na mfululizo $S_n$ una thamani

<pre>2, 4, 5, 6, 8, ...</pre>

Mfululizo $R_n$ unafafanuliwa kwa uhusiano wa urudiaji $R_n = R_{n-1} + S_{n-1}$, wakati $S_n$ unafafanuliwa kama mfululizo wa nambari kamili chanya ambazo hazijajumuishwa katika mfululizo $R_n$.

# --instructions--

Tengeneza vitendakazi viwili vinavyoitwa `ffr` na `ffs` vinavyorejesha `R(n)` au `S(n)`, mtawalia, kwa kielezo chochote `n`. Kumbuka kuwa mfululizo wa Hofstadter Figure-Figure una kielezo cha kuanzia 1, na $R_1 = 1$ na $S_1 = 2$.

Hakuna thamani kubwa ya `n` inayopaswa kudhaniwa.

**Rejea**

 <p>Rosetta: <a href='https://rosettacode.org/wiki/Hofstadter_Figure-Figure_sequences' target='_blank'>Mfululizo wa Hofstadter Figure-Figure</a></p>.


# --before-each--

```js
const ffrParamRes = [[10, 69], [50, 1509], [100, 5764], [1000, 526334]];
const ffsParamRes = [[10, 14], [50, 59], [100, 112], [1000, 1041]];
```

# --hints--

`ffr` inapaswa kuwa kitendakazi.

```js
assert(typeof ffr === 'function');
```

`ffs` inapaswa kuwa kitendakazi.

```js
assert(typeof ffs === 'function');
```

`ffr` inapaswa kurejesha nambari kamili.

```js
assert(Number.isInteger(ffr(1)));
```

`ffs` inapaswa kurejesha nambari kamili.

```js
assert(Number.isInteger(ffs(1)));
```

`ffr(10)` inapaswa kurejesha `69`

```js
assert.equal(ffr(ffrParamRes[0][0]), ffrParamRes[0][1]);
```

`ffr(50)` inapaswa kurejesha `1509`

```js
assert.equal(ffr(ffrParamRes[1][0]), ffrParamRes[1][1]);
```

`ffr(100)` inapaswa kurejesha `5764`

```js
assert.equal(ffr(ffrParamRes[2][0]), ffrParamRes[2][1]);
```

`ffr(1000)` inapaswa kurejesha `526334`

```js
assert.equal(ffr(ffrParamRes[3][0]), ffrParamRes[3][1]);
```

`ffs(10)` inapaswa kurejesha `14`

```js
assert.equal(ffs(ffsParamRes[0][0]), ffsParamRes[0][1]);
```

`ffs(50)` inapaswa kurejesha `59`

```js
assert.equal(ffs(ffsParamRes[1][0]), ffsParamRes[1][1]);
```

`ffs(100)` inapaswa kurejesha `112`

```js
assert.equal(ffs(ffsParamRes[2][0]), ffsParamRes[2][1]);
```

`ffs(1000)` inapaswa kurejesha `1041`

```js
assert.equal(ffs(ffsParamRes[3][0]), ffsParamRes[3][1]);
```

# --seed--

## --seed-contents--

```js
function ffr(n) {
  return n;
}

function ffs(n) {
  return n;
}
```

# --solutions--

```js
const R = [null, 1];
const S = [null, 2];

function extendSequences (n) {
  let current = Math.max(R[R.length - 1], S[S.length - 1]);
  let i;
  while (R.length <= n || S.length <= n) {
    i = Math.min(R.length, S.length) - 1;
    current += 1;
    if (current === R[i] + S[i]) {
      R.push(current);
    } else {
      S.push(current);
    }
  }
}

function ffr (n) {
  extendSequences(n);
  return R[n];
}

function ffs (n) {
  extendSequences(n);
  return S[n];
}
```
