---
id: 59669d08d75b60482359409f
title: Muundo wa tarehe
challengeType: 1
forumTopicId: 302243
dashedName: date-format
---

# --description--

Rudisha safu ya data yenye mfuatano wa herufi mbili za tarehe za tarehe ya sasa zenye sifa zifuatazo:

- Mfuatano wa herufi wa kwanza unapaswa kuwa na nambari ya mwaka, nambari ya mwezi, na nambari ya siku zikitenganishwa na alama za dash (`-`).
- Mwaka katika mfuatano wa herufi wa kwanza unapaswa kuwa na tarakimu nne.
- Mwezi na siku katika mfuatano wa herufi wa kwanza hazipaswi kuwa na sifuri za kuongoza.
- Jina la siku ya wiki na mwezi katika mfuatano wa herufi wa pili halipaswi kuwa kifupi.
- Siku katika mfuatano wa herufi wa pili haipaswi kuwa na sifuri za kuongoza.

Mifano ya matokeo:

```js
['2007-11-23', 'Friday, November 23, 2007']
['2021-3-2', 'Tuesday, March 2, 2021']
```

# --hints--

`getDateFormats` inapaswa kuwa kitendakazi.

```js
assert(typeof getDateFormats === 'function');
```

`getDateFormats` inapaswa kurudisha kitu.

```js
assert(typeof getDateFormats() === 'object');
```

`getDateFormats` inapaswa kurudisha safu ya data yenye vipengele 2.

```js
assert(getDateFormats().length === 2);
```

`getDateFormats` inapaswa kurudisha tarehe sahihi katika muundo unaofaa.

```js
const getDateSolution = () => {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
};

const dates = getDateSolution();
const equalsMessage = `message: <code>getDataFormats()</code> should return <code>["${dates[0]}", "${dates[1]}"]</code>.`;
assert.deepEqual(getDateFormats(), dates, equalsMessage);
```

# --seed--

## --seed-contents--

```js
function getDateFormats() {

  return true;
}
```

# --solutions--

```js
function getDateFormats() {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
}
```
