---
id: 5966f99c45e8976909a85575
title: Siku ya wiki
challengeType: 1
forumTopicId: 302245
dashedName: day-of-the-week
---

# --description--

Kampuni imeamua kwamba kila wakati Krismasi inapotokea Jumapili, watatoa likizo za ziada za kulipwa kwa wafanyakazi wao ili, pamoja na likizo za umma, wafanyakazi wasilazimike kufanya kazi wiki inayofuata (kati ya tarehe 25 Desemba na tarehe 1 Januari).

# --instructions--

Andika kitendakazi kinachopokea mwaka wa kuanzia na mwaka wa mwisho na kurudisha safu ya data ya miaka yote ambapo tarehe 25 Desemba itakuwa Jumapili.

# --before-each--

```js
const firstSolution = [1977, 1983, 1988, 1994, 2005, 2011, 2016];
const secondSolution = [2011, 2016, 2022, 2033, 2039, 2044, 2050, 2061, 2067, 2072, 2078, 2089, 2095, 2101, 2107, 2112, 2118];
```

# --hints--

`findXmasSunday` inapaswa kuwa kitendakazi.

```js
assert(typeof findXmasSunday === 'function');
```

`findXmasSunday(2000, 2100)` inapaswa kurudisha safu ya data.

```js
assert(typeof findXmasSunday(2000, 2100) === 'object');
```

`findXmasSunday(1970, 2017)` inapaswa kurudisha `[1977, 1983, 1988, 1994, 2005, 2011, 2016]`

```js
assert.deepEqual(findXmasSunday(1970, 2017), firstSolution);
```

`findXmasSunday(2008, 2121)` inapaswa kurudisha `[2011, 2016, 2022, 2033, 2039, 2044, 2050, 2061, 2067, 2072, 2078, 2089, 2095, 2101, 2107, 2112, 2118]`

```js
assert.deepEqual(findXmasSunday(2008, 2121), secondSolution);
```

# --seed--

## --seed-contents--

```js
function findXmasSunday(start, end) {

  return true;
}
```

# --solutions--

```js
function findXmasSunday(start, end) {
  const xmasSunday = [];
  for (let year = start; year <= end; year++) {
    const xmas = new Date(year, 11, 25);
    if (xmas.getDay() === 0) {
      xmasSunday.push(year);
    }
  }
  return xmasSunday;
}
```
