---
id: 59694356a6e7011f7f1c5f4e
title: Gawa kadi kwa FreeCell
challengeType: 1
forumTopicId: 302246
dashedName: deal-cards-for-freecell
---

# --description--

*FreeCell* ni mchezo wa kadi wa solitaire uliotengenezwa na Paul Alfille kwa mfumo wa PLATO mwaka 1978. Jim Horne, katika Microsoft, alibadilisha jina kuwa FreeCell na kutekeleza tena mchezo huo kwa DOS, kisha Windows. Toleo hili liliingiza gawio la kadi 32000 zilizo na nambari.

Mchezo ulipopata umaarufu, Jim Horne alifunua algoriti, na utekelezaji mwingine wa FreeCell ulianza kuiga gawio la Microsoft. Gawio hizi zina nambari kutoka 1 hadi 32000. Matoleo mapya ya Microsoft yana gawio milioni 1, zikiwa na nambari kutoka 1 hadi 1000000; baadhi ya utekelezaji huruhusu nambari nje ya anuwai hiyo.

Algoriti hutumia kizalishaji cha nambari za nasibu cha mstari (linear congruential generator) kutoka Microsoft C:

<ul>
  <li>$state_{n + 1} \equiv 214013 \times state_n + 2531011 \pmod{2^{31}}$</li>
  <li>$rand_n = state_n \div 2^{16}$</li>
  <li>$rand_n$ iko katika anuwai 0 hadi 32767.</li>
</ul>

Algoriti inafuata hatua hizi:

<ol>
  <li>Weka mbegu (seed) ya RNG kwa nambari ya gawio.
  </li><li>Tengeneza safu ya data ya kadi 52: Ace wa Clubs, Ace wa Diamonds, Ace wa Hearts, Ace wa Spades, 2 wa Clubs, 2 wa Diamonds, na kadhalika kupitia viwango: Ace, 2, 3, 4, 5, 6, 7, 8, 9, 10, Jack, Queen, King. Kielezo cha safu ni 0 hadi 51, na Ace wa Clubs iko 0, na King wa Spades iko 51.</li>
  <li>Mpaka safu ya data iwe tupu:</li>
    <ul>
      <li>Chagua kadi ya nasibu kwa <i>kielezo</i> ≡ <i>nambari inayofuata ya nasibu</i> (mod <i>urefu wa safu ya data</i>).</li>
      <li>Badilisha kadi hii ya nasibu na kadi ya mwisho ya safu ya data.</li>
      <li>Ondoa kadi hii ya nasibu kutoka safu ya data. (Urefu wa safu hupungua kwa 1.)</li>
      <li>Gawa kadi hii ya nasibu.</li>
    </ul>
  <li>Gawa kadi zote 52, uso juu, katika nguzo 8. Kadi 8 za kwanza zinaenda nguzo 8, kadi 8 zinazofuata zinaenda juu ya kadi 8 za kwanza, na kadhalika.</li>
</ol>

**Mfano:**

**Mpangilio wa kugawa kadi**

<pre> 1  2  3  4  5  6  7  8
 9 10 11 12 13 14 15 16
17 18 19 20 21 22 23 24
25 26 27 28 29 30 31 32
33 34 35 36 37 38 39 40
41 42 43 44 45 46 47 48
49 50 51 52</pre>

**Mchezo #1**

```js
[
  ['JD', '2D', '9H', 'JC', '5D', '7H', '7C', '5H'],
  ['KD', 'KC', '9S', '5S', 'AD', 'QC', 'KH', '3H'],
  ['2S', 'KS', '9D', 'QD', 'JS', 'AS', 'AH', '3C'],
  ['4C', '5C', 'TS', 'QH', '4H', 'AC', '4D', '7S'],
  ['3S', 'TD', '4S', 'TH', '8H', '2C', 'JH', '7D'],
  ['6D', '8S', '8D', 'QS', '6C', '3D', '8C', 'TC'],
  ['6S', '9C', '2H', '6H']
]
```

**Mchezo #617**

```js
[
  ['7D', 'AD', '5C', '3S', '5S', '8C', '2D', 'AH'],
  ['TD', '7S', 'QD', 'AC', '6D', '8H', 'AS', 'KH'],
  ['TH', 'QC', '3H', '9D', '6S', '8D', '3D', 'TC'],
  ['KD', '5H', '9S', '3C', '8S', '7H', '4D', 'JS'],
  ['4C', 'QS', '9C', '9H', '7C', '6H', '2C', '2S'],
  ['4S', 'TS', '2H', '5D', 'JC', '6C', 'JH', 'QH'],
  ['JD', 'KS', 'KC', '4H']
]
```

# --instructions--

Andika kitendakazi kinachopokea nambari ya gawio na kugawa kadi kwa mpangilio sawa na algoriti hii. Kitendakazi lazima kirudishe safu ya data ya vipimo viwili inayo wakilisha ubao wa FreeCell.

# --before-each--

```js
const replaceThis = 3;
const game1 = [
  ['JD', '2D', '9H', 'JC', '5D', '7H', '7C', '5H'],
  ['KD', 'KC', '9S', '5S', 'AD', 'QC', 'KH', '3H'],
  ['2S', 'KS', '9D', 'QD', 'JS', 'AS', 'AH', '3C'],
  ['4C', '5C', 'TS', 'QH', '4H', 'AC', '4D', '7S'],
  ['3S', 'TD', '4S', 'TH', '8H', '2C', 'JH', '7D'],
  ['6D', '8S', '8D', 'QS', '6C', '3D', '8C', 'TC'],
  ['6S', '9C', '2H', '6H']
];
const game617 = [
  ['7D', 'AD', '5C', '3S', '5S', '8C', '2D', 'AH'],
  ['TD', '7S', 'QD', 'AC', '6D', '8H', 'AS', 'KH'],
  ['TH', 'QC', '3H', '9D', '6S', '8D', '3D', 'TC'],
  ['KD', '5H', '9S', '3C', '8S', '7H', '4D', 'JS'],
  ['4C', 'QS', '9C', '9H', '7C', '6H', '2C', '2S'],
  ['4S', 'TS', '2H', '5D', 'JC', '6C', 'JH', 'QH'],
  ['JD', 'KS', 'KC', '4H']
];
```

# --hints--

`dealFreeCell` inapaswa kuwa kitendakazi.

```js
assert(typeof dealFreeCell === 'function');
```

`dealFreeCell(seed)` inapaswa kurudisha kitu.

```js
assert(typeof dealFreeCell(1) === 'object');
```

`dealFreeCell(seed)` inapaswa kurudisha safu ya data yenye urefu wa 7.

```js
assert(dealFreeCell(1).length === 7);
```

`dealFreeCell(1)` inapaswa kurudisha safu ya data sawa na mfano "Mchezo #1"

```js
assert.deepEqual(dealFreeCell(1), game1);
```

`dealFreeCell(617)` inapaswa kurudisha safu ya data sawa na mfano "Mchezo #617"

```js
assert.deepEqual(dealFreeCell(617), game617);
```

# --seed--

## --seed-contents--

```js
function dealFreeCell(seed) {

  return true;
}
```

# --solutions--

```js
// RNG
function FreeCellRNG(seed) {
  return {
    lastNum: seed,
    next() {
      this.lastNum = ((214013 * this.lastNum) + 2531011) % (Math.pow(2, 31));
      return this.lastNum >> 16;
    }
  };
}
// Get cards
function getDeck() {
  const ranks = ['A', '2', '3', '4', '5', '6', '7', '8', '9', 'T', 'J', 'Q', 'K'];
  const suits = ['C', 'D', 'H', 'S'];
  const cards = [];
  for (let i = 0; i < ranks.length; i += 1) {
    for (let j = 0; j < suits.length; j += 1) {
      cards.push(`${ranks[i]}${suits[j]}`);
    }
  }
  return cards;
}
function dealFreeCell(seed) {
  const rng = FreeCellRNG(seed);
  const deck = getDeck();

  const dealtCards = [[], [], [], [], [], [], []];
  let currentColumn = 0;
  let currentRow = 0;

  let rand;
  let temp;
  let card;
  while (deck.length > 0) {
    // Choose a random card
    rand = rng.next() % deck.length;

    // Swap this random card with the last card in the array
    temp = deck[deck.length - 1];
    deck[deck.length - 1] = deck[rand];
    deck[rand] = temp;

    // Remove this card from the array
    card = deck.pop();

    // Deal this card
    dealtCards[currentRow].push(card);
    currentColumn += 1;
    if (currentColumn === 8) {
      currentColumn = 0;
      currentRow += 1;
    }
  }

  return dealtCards;
}
```
