---
id: 596a8888ab7c01048de257d5
title: Nakala ya kina
challengeType: 1
forumTopicId: 302247
dashedName: deepcopy
---

# --description--

Andika kitendakazi kinachorejesha nakala ya kina ya kitu kilichotolewa. Nakala hiyo haipaswi kuwa kitu kilekile kilichotolewa.

Zoezi hili halitahakikisha:

<ul>
  <li>Vitu vyenye vigezo ambavyo ni vitendakazi</li>
  <li>Vitu vya tarehe au vitu vyenye vigezo ambavyo ni vitu vya tarehe</li>
  <li>Usemi wa kawaida au vitu vyenye vigezo ambavyo ni vitu vya usemi wa kawaida</li>
  <li>Unakili wa mfano</li>
</ul>

# --before-each--

```js
const obj1 = { test: 'test' };
const obj2 = {
  t: 'test',
  a: ['an', 'array']
};
const obj3 = {
  t: 'try',
  o: obj2
};
```

# --hints--

`deepcopy` inapaswa kuwa kitendakazi.

```js
assert(typeof deepcopy === 'function');
```

`deepcopy({test: "test"})` inapaswa kurejesha kitu.

```js
assert(typeof deepcopy(obj1) === 'object');
```

`deepcopy` haipaswi kurejesha kitu kilekile kilichotolewa.

```js
assert(deepcopy(obj2) != obj2);
```

Unapopokea kitu chenye safu ya data, `deepcopy` inapaswa kurejesha nakala ya kina ya kitu hicho.

```js
assert.deepEqual(deepcopy(obj2), obj2);
```

Unapopokea kitu chenye kitu kingine ndani yake, `deepcopy` inapaswa kurejesha nakala ya kina ya kitu hicho.

```js
assert.deepEqual(deepcopy(obj3), obj3);
```

# --seed--

## --seed-contents--

```js
function deepcopy(obj) {

  return true;
}
```

# --solutions--

```js
function deepcopy(obj) {
  return JSON.parse(JSON.stringify(obj));
}
```
