---
id: 596e414344c3b2872167f0fe
title: Mjadala wa koma
challengeType: 1
forumTopicId: 302234
dashedName: comma-quibbling
---

# --description--

<a href="https://rosettacode.org/wiki/Comma_quibbling" target="_blank" rel="noopener noreferrer nofollow">Mjadala wa koma</a> ni zoezi lililotanguliwa na Eric Lippert katika blogu yake.

# --instructions--

Andika kitendakazi cha kuzalisha mfuatano wa herufi unaotokana na kuunganisha maneno ya ingizo kutoka kwenye orodha/mfululizo ambapo:

<ol>
  <li>Ingizo lisilo na maneno lolote hutoa mfuatano wa herufi wa herufi mbili za mabano (<code>"{}"</code>)</li>
  <li>Ingizo la neno moja tu, mfano <code>["ABC"]</code>, hutoa mfuatano wa herufi wa neno hilo ndani ya mabano hayo mawili, mfano <code>"{ABC}"</code></li>
  <li>Ingizo la maneno mawili, mfano <code>["ABC", "DEF"]</code>, hutoa mfuatano wa herufi wa maneno hayo mawili ndani ya mabano mawili yakiwa yamegawanywa na mfuatano wa herufi <code>" and "</code>, mfano <code>"{ABC and DEF}"</code></li>
  <li>Ingizo la maneno matatu au zaidi, mfano <code>["ABC", "DEF", "G", "H"]</code>, hutoa mfuatano wa herufi wa maneno yote isipokuwa neno la mwisho yakiwa yamegawanywa na <code>", "</code> na neno la mwisho likiwa limegawanywa na <code>" and "</code> na yote yakiwa ndani ya mabano; mfano <code>"{ABC, DEF, G and H}"</code></li>
</ol>

Jaribu kitendakazi chako kwa mfululizo wa ingizo uonyeshe matokeo yako hapa kwenye ukurasa huu:

<ul>
  <li>[] # (Hakuna maneno ya ingizo).</li>
  <li>["ABC"]</li>
  <li>["ABC", "DEF"]</li>
  <li>["ABC", "DEF", "G", "H"]</li>
</ul>

**Kumbuka:** Kubali maneno kuwa mfuatano wa herufi usio na tupu wa herufi kubwa kwa zoezi hili.

# --before-each--

```js
const testCases = [[], ["ABC"], ["ABC", "DEF"], ["ABC", "DEF", "G", "H"]];
const results = ["{}", "{ABC}", "{ABC and DEF}", "{ABC, DEF, G and H}"];
```

# --hints--

`quibble` inapaswa kuwa kitendakazi.

```js
assert(typeof quibble === 'function');
```

`quibble(["ABC"])` inapaswa kurudisha mfuatano wa herufi.

```js
assert(typeof quibble(['ABC']) === 'string');
```

`quibble([])` inapaswa kurudisha "{}".

```js
assert.equal(quibble(testCases[0]), results[0]);
```

`quibble(["ABC"])` inapaswa kurudisha `"{ABC}"`.

```js
assert.equal(quibble(testCases[1]), results[1]);
```

`quibble(["ABC", "DEF"])` inapaswa kurudisha `"{ABC and DEF}"`.

```js
assert.equal(quibble(testCases[2]), results[2]);
```

`quibble(["ABC", "DEF", "G", "H"])` inapaswa kurudisha `"{ABC, DEF, G and H}"`.

```js
assert.equal(quibble(testCases[3]), results[3]);
```

# --seed--

## --seed-contents--

```js
function quibble(words) {

  return true;
}
```

# --solutions--

```js
function quibble(words) {
  return "{" +
    words.slice(0, words.length - 1).join(", ") +
   (words.length > 1 ? " and " : "") +
   (words[words.length - 1] || '') +
  "}";
}
```
