---
id: 596e457071c35c882915b3e4
title: Linganisha orodha ya mfuatano wa herufi
challengeType: 1
forumTopicId: 302235
dashedName: compare-a-list-of-strings
---

# --description--

Orodha ni seti iliyopangwa ya maadili ambayo inaweza kuwa na rudufu. Hapa kuna mfano:

```js
const list = [['AA',  'BB', 'CC'], ['AA', 'ACB', 'AA'], [], ['AA']];
````

Ikiwa umepewa orodha ya mfuatano wa herufi kwa idadi yoyote ile, tekeleza kitendakazi kwa kila mojawapo ya masharti yafuatayo:

<ul>
  <li>jaribu kama zote ni sawa kwa mtindo wa kihesabu wa mfuatano wa herufi</li>
  <li>jaribu kama kila mfuatano wa herufi ni mdogo kihesabu kuliko ule unaofuata (yaani kama orodha iko katika mpangilio mkali wa kuongezeka)</li>
</ul>

# --before-each--

```js
const testCases = [['AA', 'AA', 'AA', 'AA'], ['AA', 'ACB', 'BB', 'CC'], [], ['AA'], ['BB', 'AA']];
```

# --hints--

`allEqual` inapaswa kuwa kitendakazi.

```js
assert(typeof allEqual === 'function');
```

`azSorted` inapaswa kuwa kitendakazi.

```js
assert(typeof azSorted === 'function');
```

`allEqual(["AA", "AA", "AA", "AA"])` inapaswa kurudisha true.

```js
assert(allEqual(testCases[0]));
```

`azSorted(["AA", "AA", "AA", "AA"])` inapaswa kurudisha false.

```js
assert(!azSorted(testCases[0]));
```

`allEqual(["AA", "ACB", "BB", "CC"])` inapaswa kurudisha false.

```js
assert(!allEqual(testCases[1]));
```

`azSorted(["AA", "ACB", "BB", "CC"])` inapaswa kurudisha true.

```js
assert(azSorted(testCases[1]));
```

`allEqual([])` inapaswa kurudisha true.

```js
assert(allEqual(testCases[2]));
```

`azSorted([])` inapaswa kurudisha true.

```js
assert(azSorted(testCases[2]));
```

`allEqual(["AA"])` inapaswa kurudisha true.

```js
assert(allEqual(testCases[3]));
```

`azSorted(["AA"])` inapaswa kurudisha true.

```js
assert(azSorted(testCases[3]));
```

`allEqual(["BB", "AA"])` inapaswa kurudisha false.

```js
assert(!allEqual(testCases[4]));
```

`azSorted(["BB", "AA"])` inapaswa kurudisha false.

```js
assert(!azSorted(testCases[4]));
```

# --seed--

## --seed-contents--

```js
function allEqual(arr) {

  return true;
}

function azSorted(arr) {

  return true;
}
```

# --solutions--

```js
function allEqual(a) {
  let out = true;
  let i = 0;
  while (++i < a.length) {
    out = out && (a[i - 1] === a[i]);
  } return out;
}

function azSorted(a) {
  let out = true;
  let i = 0;
  while (++i < a.length) {
    out = out && (a[i - 1] < a[i]);
  } return out;
}
```
