---
id: 597089c87eec450c68aa1643
title: Eleza aina ya data ya msingi
challengeType: 1
forumTopicId: 302248
dashedName: define-a-primitive-data-type
---

# --description--

Eleza aina inayofanya kazi kama nambari kamili lakini ina thamani halali ya chini kabisa ya 1 na thamani halali ya juu kabisa ya 10.

Usimamizi wa makosa:

<ul>
  <li>Kama utajaribu kuanzisha <code>Num</code> na thamani nje ya 1 - 10, inapaswa tupa <code>TypeError</code> yenye ujumbe wa makosa <code>'Out of range'</code>.</li>
  <li>Kama utajaribu kuanzisha <code>Num</code> na thamani ambayo si nambari, inapaswa tupa <code>TypeError</code> yenye ujumbe wa makosa <code>'Not a Number'</code>.</li>
</ul>

# --hints--

`Num` inapaswa kuwa kitendakazi.

```js
assert(typeof Num === 'function');
```

`new Num(4)` inapaswa kurudisha kitu.

```js
assert(typeof new Num(4) === 'object');
```

`new Num('test')` inapaswa kutupa TypeError yenye ujumbe 'Not a Number'.

```js
assert.throws(() => new Num('test'), TypeError);
```

`new Num(0)` inapaswa kutupa TypeError yenye ujumbe 'Out of range'.

```js
assert.throws(() => new Num(0), TypeError);
```

`new Num(-5)` inapaswa kutupa TypeError yenye ujumbe 'Out of range'.

```js
assert.throws(() => new Num(-5), TypeError);
```

`new Num(10)` inapaswa kutupa TypeError yenye ujumbe 'Out of range'.

```js
assert.throws(() => new Num(11), TypeError);
```

`new Num(20)` inapaswa kutupa TypeError yenye ujumbe 'Out of range'.

```js
assert.throws(() => new Num(20), TypeError);
```

`new Num(3) + new Num(4)` inapaswa kuwa sawa na 7.

```js
assert.equal(new Num(3) + new Num(4), 7);
```

`new Num(3) - new Num(4)` inapaswa kuwa sawa na -1.

```js
assert.equal(new Num(3) - new Num(4), -1);
```

`new Num(3) * new Num(4)` inapaswa kuwa sawa na 12.

```js
assert.equal(new Num(3) * new Num(4), 12);
```

`new Num(3) / new Num(4)` inapaswa kuwa sawa na 0.75.

```js
assert.equal(new Num(3) / new Num(4), 0.75);
```

`new Num(3) < new Num(4)` inapaswa kuwa kweli.

```js
assert(new Num(3) < new Num(4));
```

`new Num(3) > new Num(4)` inapaswa kuwa si kweli.

```js
assert(!(new Num(3) > new Num(4)));
```

`(new Num(5)).toString()` inapaswa kurudisha '5'

```js
assert.equal(new Num(5).toString(), '5');
```

# --seed--

## --seed-contents--

```js
function Num(n) {

  return n;
}
```

# --solutions--

```js
function Num(n) {
  if (isNaN(n)) {
    throw new TypeError('Not a Number');
  }
  if (n < 1 || n > 10) {
    throw new TypeError('Out of range');
  }

  this._value = +n;
}
Num.prototype.valueOf = function() { return this._value; };
Num.prototype.toString = function () { return this._value.toString(); };
```
