---
id: 59713bd26bdeb8a594fb9413
title: Hesabu sarafu
challengeType: 1
forumTopicId: 302238
dashedName: count-the-coins
---

# --description--

Kuna aina nne za sarafu zinazotumika sana katika sarafu za Marekani:

<ul>
  <li>quarters (25 senti)</li>
  <li>dimes (10 senti)</li>
  <li>nickels (5 senti), na</li>
  <li>pennies (1 senti)</li>
</ul>

<p>Kuna njia sita za kutoa mabadiliko ya senti 15:</p>

<ul>
  <li>Dime na nickel</li>
  <li>Dime na pennies 5</li>
  <li>nickels 3</li>
  <li>nickels 2 na pennies 5</li>
  <li>nickel 1 na pennies 10</li>
  <li>pennies 15</li>
</ul>

# --instructions--

Tekeleza kitendakazi ili kubaini ni njia ngapi zinapatikana za kutoa mabadiliko kwa ingizo fulani, `cents`, kinachowakilisha kiasi cha senti za Marekani kwa kutumia sarafu hizi zinazotumika sana.

# --hints--

`countCoins` inapaswa kuwa kitendakazi.

```js
assert(typeof countCoins === 'function');
```

`countCoins(15)` inapaswa kurudisha `6`.

```js
assert.equal(countCoins(15), 6);
```

`countCoins(85)` inapaswa kurudisha `163`.

```js
assert.equal(countCoins(85), 163);
```

`countCoins(100)` inapaswa kurudisha `242`.

```js
assert.equal(countCoins(100), 242);
```

# --seed--

## --seed-contents--

```js
function countCoins(cents) {

  return true;
}
```

# --solutions--

```js
function countCoins(cents) {
  const operands = [1, 5, 10, 25];
  const targetsLength = cents + 1;
  const operandsLength = operands.length;
  const t = [1];

  for (let a = 0; a < operandsLength; a++) {
    for (let b = 1; b < targetsLength; b++) {
      // initialise undefined target
      t[b] = t[b] ? t[b] : 0;

      // accumulate target + operand ways
      t[b] += (b < operands[a]) ? 0 : t[b - operands[a]];
    }
  }

  return t[targetsLength - 1];
}
```
