---
id: 597f1e7fbc206f0e9ba95dc4
title: Vizingiti vya nambari kamili
challengeType: 1
forumTopicId: 302265
dashedName: factors-of-an-integer
---

# --description--

Andika kitendakazi kinachorejesha vizingiti vya nambari kamili chanya kama safu ya data.

Vizingiti hivi ni nambari kamili chanya ambavyo nambari inayotafutwa vizingiti inaweza kugawanywa navyo ili kupata matokeo ya nambari kamili chanya.

# --before-each--

```js
const ans=[[1,3,5,9,15,45],[1,53],[1,2,4,8,16,32,64]];
```

# --hints--

`factors` inapaswa kuwa kitendakazi.

```js
assert(typeof factors === 'function');
```

`factors(45)` inapaswa kurejesha `[1,3,5,9,15,45]`.

```js
assert.deepEqual(factors(45), ans[0]);
```

`factors(53)` inapaswa kurejesha `[1,53]`.

```js
assert.deepEqual(factors(53), ans[1]);
```

`factors(64)` inapaswa kurejesha `[1,2,4,8,16,32,64]`.

```js
assert.deepEqual(factors(64), ans[2]);
```

# --seed--

## --seed-contents--

```js
function factors(num) {

}
```

# --solutions--

```js
function factors(num)
{
 let n_factors = [], i, sqr=Math.floor(Math.sqrt(num));

 for (i = 1; i <=sqr ; i += 1)
  if (num % i === 0)
  {
   n_factors.push(i);
   if (num / i !== i)
    n_factors.push(num / i);
  }
 n_factors.sort(function(a, b){return a - b;});
 return n_factors;
}
```
