---
id: 597f24c1dda4e70f53c79c81
title: Mfululizo wa Fibonacci
challengeType: 1
forumTopicId: 302268
dashedName: fibonacci-sequence
---

# --description--

Andika kitendakazi cha kuzalisha nambari ya nambari ya <code>n<sup>th</sup></code> katika mfululizo wa Fibonacci.

Nambari ya <code>n<sup>th</sup></code> katika mfululizo wa Fibonacci hutolewa na:

<code>F<sub>n</sub> = F<sub>n-1</sub> + F<sub>n-2</sub></code>

Vipengele viwili vya kwanza vya mfululizo ni 0 na 1.

Hivyo basi, mfululizo ni: 0, 1, 1, 2, 3, 5, 8, 13...

# --hints--

`fibonacci` inapaswa kuwa kitendakazi.

```js
assert(typeof fibonacci === 'function');
```

`fibonacci(2)` inapaswa kurudisha nambari.

```js
assert(typeof fibonacci(2) == 'number');
```

`fibonacci(3)` inapaswa kurudisha 2.

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` inapaswa kurudisha 5.

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` inapaswa kurudisha 55.

```js
assert.equal(fibonacci(10), 55);
```

# --seed--

## --seed-contents--

```js
function fibonacci(n) {

}
```

# --solutions--

```js
function fibonacci(n) {
  let a = 0, b = 1, t;
  while (--n >= 0) {
    t = a;
    a = b;
    b += t;
  }
  return a;
}
```
