---
id: 59880443fb36441083c6c20e
title: Njia ya Euler
challengeType: 1
forumTopicId: 302258
dashedName: euler-method
---

# --description--

Njia ya Euler inakisia kwa nambari suluhisho za mlinganyo wa tofauti wa kawaida wa daraja la kwanza (ODEs) kwa thamani ya awali iliyotolewa. Ni njia wazi ya kutatua matatizo ya thamani ya awali (IVPs), kama ilivyoelezwa katika <a href="https://www.freecodecamp.org/news/eulers-method-explained-with-examples/" title="Euler's Method Explained with Examples" target="_blank" rel="noopener noreferrer nofollow">makala hii</a>.

ODE lazima itolewe kwa muundo ufuatao:

<ul style='list-style: none;'>
  <li><big>$\frac{dy(t)}{dt} = f(t,y(t))$</big></li>
</ul>

pamoja na thamani ya awali

<ul style='list-style: none;'>
  <li><big>$y(t_0) = y_0$</big></li>
</ul>

Ili kupata suluhisho la nambari, tunabadilisha mtoaji upande wa kushoto (LHS) kwa makisio ya tofauti ya mwisho:

<ul style='list-style: none;'>
  <li><big>$\frac{dy(t)}{dt}  \approx \frac{y(t+h)-y(t)}{h}$</big></li>
</ul>

kisha tatua kwa $y(t+h)$:

<ul style='list-style: none;'>
  <li><big>$y(t+h) \approx y(t) + h \, \frac{dy(t)}{dt}$</big></li>
</ul>

ambayo ni sawa na

<ul style='list-style: none;'>
  <li><big>$y(t+h) \approx y(t) + h \, f(t,y(t))$</big></li>
</ul>

Kanuni ya suluhisho ya kurudia ni:

<ul style='list-style: none;'>
  <li><big>$y_{n+1} = y_n + h \, f(t_n, y_n)$</big></li>
</ul>

ambapo $h$ ni ukubwa wa hatua, kigezo muhimu zaidi kwa usahihi wa suluhisho. Ukubwa mdogo wa hatua huongeza usahihi lakini pia gharama ya mahesabu, hivyo daima lazima uchaguliwe kwa mkono kulingana na changamoto iliyopo.

**Mfano: Sheria ya kupoeza ya Newton**

Sheria ya kupoeza ya Newton inaelezea jinsi kitu chenye joto la awali $T(t_0) = T_0$ kinapoezwa katika mazingira yenye joto $T_R$:

<ul style='list-style: none;'>
  <li><big>$\frac{dT(t)}{dt} = -k \, \Delta T$</big></li>
</ul>

au

<ul style='list-style: none;'>
  <li><big>$\frac{dT(t)}{dt} = -k \, (T(t) - T_R)$</big></li>
</ul>

Inasema kuwa kiwango cha kupoeza $\\frac{dT(t)}{dt}$ cha kitu hicho kinahusiana na tofauti ya sasa ya joto $\\Delta T = (T(t) - T_R)$ na mazingira yanayozunguka.

Suluhisho la kihesabu, ambalo tutalinganisha na makisio ya nambari, ni

<ul style='list-style: none;'>
  <li><big>$T(t) = T_R + (T_0 - T_R) \; e^{-k t}$</big></li>
</ul>

# --instructions--

Tekeleza utaratibu wa njia ya Euler kisha utumie kutatua mfano uliotolewa wa sheria ya kupoeza ya Newton kwa ukubwa wa hatua tatu tofauti za:

<ul>
  <li><code>2 s</code></li>
  <li><code>5 s</code> na</li>
  <li><code>10 s</code></li>
</ul>

kisha linganisha na suluhisho la kihesabu.

**Thamani za awali:**

<ul>
  <li>joto la awali <big>$T_0$</big> litakuwa <code>100 °C</code></li>
  <li>joto la chumba <big>$T_R$</big> litakuwa <code>20 °C</code></li>
  <li>konstanti ya kupoeza <big>$k$</big> itakuwa <code>0.07</code></li>
  <li>kipindi cha wakati cha kuhesabu kitakuwa kutoka <code>0 s</code> hadi <code>100 s</code></li>
</ul>  

Kigezo cha kwanza cha kitendakazi ni wakati wa awali, kigezo cha pili ni joto la awali, kigezo cha tatu ni muda uliopita na kigezo cha nne ni ukubwa wa hatua.

# --hints--

`eulersMethod` inapaswa kuwa kitendakazi.

```js
assert(typeof eulersMethod === 'function');
```

`eulersMethod(0, 100, 100, 2)` inapaswa kurudisha nambari.

```js
assert(typeof eulersMethod(0, 100, 100, 2) === 'number');
```

`eulersMethod(0, 100, 100, 2)` inapaswa kurudisha 20.0424631833732.

```js
assert.equal(eulersMethod(0, 100, 100, 2), 20.0424631833732);
```

`eulersMethod(0, 100, 100, 5)` inapaswa kurudisha 20.01449963666907.

```js
assert.equal(eulersMethod(0, 100, 100, 5), 20.01449963666907);
```

`eulersMethod(0, 100, 100, 10)` inapaswa kurudisha 20.000472392.

```js
assert.equal(eulersMethod(0, 100, 100, 10), 20.000472392);
```

# --seed--

## --seed-contents--

```js
function eulersMethod(x1, y1, x2, h) {

}
```

# --solutions--

```js
function eulersMethod(x1, y1, x2, h) {
  let x = x1;
  let y = y1;

  while ((x < x2 && x1 < x2) || (x > x2 && x1 > x2)) {
    y += h * (-0.07 * (y - 20));
    x += h;
  }

  return y;
}
```
