---
id: 598de241872ef8353c58a7a2
title: Hesabu viwango vya binomial
challengeType: 1
forumTopicId: 302259
dashedName: evaluate-binomial-coefficients
---

# --description--

Andika kitendakazi cha kuhesabu kiwango cha binomial kwa thamani iliyotolewa ya n na k.

Fomula hii inapendekezwa:

$\\binom{n}{k} = \\frac{n!}{(n-k)!k!} = \\frac{n(n-1)(n-2)\\ldots(n-k+1)}{k(k-1)(k-2)\\ldots 1}$

# --hints--

`binom` inapaswa kuwa kitendakazi.

```js
assert(typeof binom === 'function');
```

`binom(5,3)` inapaswa kurudisha 10.

```js
assert.equal(binom(5, 3), 10);
```

`binom(7,2)` inapaswa kurudisha 21.

```js
assert.equal(binom(7, 2), 21);
```

`binom(10,4)` inapaswa kurudisha 210.

```js
assert.equal(binom(10, 4), 210);
```

`binom(6,1)` inapaswa kurudisha 6.

```js
assert.equal(binom(6, 1), 6);
```

`binom(12,8)` inapaswa kurudisha 495.

```js
assert.equal(binom(12, 8), 495);
```

# --seed--

## --seed-contents--

```js
function binom(n, k) {

}
```

# --solutions--

```js
function binom(n, k) {
  let coeff = 1;
  for (let i = n - k + 1; i <= n; i++) coeff *= i;
  for (let i = 1; i <= k; i++) coeff /= i;
  return coeff;
}
```
