---
id: 598eef80ba501f1268170e1e
title: Mfululizo wa nambari wa Fibonacci wa hatua n
challengeType: 1
forumTopicId: 302267
dashedName: fibonacci-n-step-number-sequences
---

# --description--

Mfululizo huu wa nambari ni upanuzi wa mfululizo wa kawaida wa Fibonacci ambapo:

<ol>
  <li>Kwa $n = 2$ tunapata mfululizo wa Fibonacci; wenye thamani za awali $[1, 1]$ na $F_k^2 = F_{k-1}^2 + F_{k-2}^2$</li>
  <li>Kwa $n = 3$ tunapata mfululizo wa tribonacci; wenye thamani za awali $[1, 1, 2]$ na $F_k^3 = F_{k-1}^3 + F_{k-2}^3 + F_{k-3}^3$</li>
  <li>Kwa $n = 4$ tunapata mfululizo wa tetranacci; wenye thamani za awali $[1, 1, 2, 4]$ na $F_k^4 = F_{k-1}^4 + F_{k-2}^4 + F_{k-3}^4 + F_{k-4}^4$...</li>
  <li>Kwa $n>2$ kwa ujumla tunapata mfululizo wa hatua n wa Fibonacci - $F_k^n$; wenye thamani za awali za nambari $n$ za kwanza za mfululizo wa hatua $(n-1)$ wa Fibonacci $F_k^{n-1}$; na thamani ya $k$ ya mfululizo huu wa $n$ ni $F_k^n = \sum_{i=1}^{(n)} {F_{k-i}^{(n)}}$</li>
</ol>

Kwa thamani ndogo za $n$, mara nyingine hutumika kiambishi awali cha nambari za Kigiriki kwa kupewa majina binafsi kila mfululizo.

Mfululizo wa hatua n wa Fibonacci:

| $n$ | Jina la mfululizo | Thamani                                                 |
| --- | ----------------- | ------------------------------------------------------- |
| 2   | fibonacci         | 1 1 2 3 5 8 13 21 34 55 89 144 233 377 610 ...          |
| 3   | tribonacci        | 1 1 2 4 7 13 24 44 81 149 274 504 927 1705 3136 ...     |
| 4   | tetranacci        | 1 1 2 4 8 15 29 56 108 208 401 773 1490 2872 5536 ...   |
| 5   | pentanacci        | 1 1 2 4 8 16 31 61 120 236 464 912 1793 3525 6930 ...   |
| 6   | hexanacci         | 1 1 2 4 8 16 32 63 125 248 492 976 1936 3840 7617 ...   |
| 7   | heptanacci        | 1 1 2 4 8 16 32 64 127 253 504 1004 2000 3984 7936 ...  |
| 8   | octonacci         | 1 1 2 4 8 16 32 64 128 255 509 1016 2028 4048 8080 ...  |
| 9   | nonanacci         | 1 1 2 4 8 16 32 64 128 256 511 1021 2040 4076 8144 ...  |
| 10  | decanacci         | 1 1 2 4 8 16 32 64 128 256 512 1023 2045 4088 8172 ...  |

Mfululizo unaohusiana unaweza kuzalishwa ambapo thamani za awali zinabadilishwa: Mfululizo wa Lucas hujumlisha thamani mbili zilizotangulia kama mfululizo wa Fibonacci kwa $n=2$ lakini hutumia $\[2, 1]$ kama thamani zake za awali.

# --instructions--

Andika kitendakazi kuzalisha mfululizo wa nambari wa hatua n wa Fibonacci na mfululizo wa Lucas. Kigezo cha kwanza kitakuwa $n$. Kigezo cha pili kitakuwa idadi ya vipengele vinavyorudishwa. Kigezo cha tatu kitaeleza kama utatolewa mfululizo wa Fibonacci au mfululizo wa Lucas. Ikiwa kigezo ni `"f"` basi rudisha mfululizo wa Fibonacci na ikiwa ni `"l"` basi rudisha mfululizo wa Lucas. Mfululizo lazima urudishwe kama safu ya data.

# --hints--

`fib_luc` inapaswa kuwa kitendakazi.

```js
assert(typeof fib_luc === 'function');
```

`fib_luc(2,10,"f")` inapaswa kurudisha `[1,1,2,3,5,8,13,21,34,55]`.

```js
assert.deepEqual(fib_luc(2, 10, 'f'), globalThis.ans[0]);
```

`fib_luc(3,15,"f")` inapaswa kurudisha `[1,1,2,4,7,13,24,44,81,149,274,504,927,1705,3136]`.

```js
assert.deepEqual(fib_luc(3, 15, 'f'), globalThis.ans[1]);
```

`fib_luc(4,15,"f")` inapaswa kurudisha `[1,1,2,4,8,15,29,56,108,208,401,773,1490,2872,5536]`.

```js
assert.deepEqual(fib_luc(4, 15, 'f'), globalThis.ans[2]);
```

`fib_luc(2,10,"l")` inapaswa kurudisha `[ 2, 1, 3, 4, 7, 11, 18, 29, 47, 76]`.

```js
assert.deepEqual(fib_luc(2, 10, 'l'), globalThis.ans[3]);
```

`fib_luc(3,15,"l")` inapaswa kurudisha `[ 2, 1, 3, 6, 10, 19, 35, 64, 118, 217, 399, 734, 1350, 2483, 4567 ]`.

```js
assert.deepEqual(fib_luc(3, 15, 'l'), globalThis.ans[4]);
```

`fib_luc(4,15,"l")` inapaswa kurudisha `[ 2, 1, 3, 6, 12, 22, 43, 83, 160, 308, 594, 1145, 2207, 4254, 8200 ]`.

```js
assert.deepEqual(fib_luc(4, 15, 'l'), globalThis.ans[5]);
```

`fib_luc(5,15,"l")` inapaswa kurudisha `[ 2, 1, 3, 6, 12, 24, 46, 91, 179, 352, 692, 1360, 2674, 5257, 10335 ]`.

```js
assert.deepEqual(fib_luc(5, 15, 'l'), globalThis.ans[6]);
```

# --before-all--

```js
globalThis.ans = [[1,1,2,3,5,8,13,21,34,55],
[1,1,2,4,7,13,24,44,81,149,274,504,927,1705,3136],
[1,1,2,4,8,15,29,56,108,208,401,773,1490,2872,5536],
[ 2, 1, 3, 4, 7, 11, 18, 29, 47, 76],
[ 2, 1, 3, 6, 10, 19, 35, 64, 118, 217, 399, 734, 1350, 2483, 4567 ],
[ 2, 1, 3, 6, 12, 22, 43, 83, 160, 308, 594, 1145, 2207, 4254, 8200 ],
[ 2, 1, 3, 6, 12, 24, 46, 91, 179, 352, 692, 1360, 2674, 5257, 10335 ]];
```

# --seed--

## --seed-contents--

```js
function fib_luc(n, len, w) {

}
```

# --solutions--

```js
function fib_luc(n, len, w) {
    function nacci(a, n, len) {
        while (a.length < len) {
            let sum = 0;
            for (let i = Math.max(0, a.length - n); i < a.length; i++)
                sum += a[i];
            a.push(sum);
        }
        return a;
    }
    if(w=="f"){
        return nacci(nacci([1,1], n, n), n, len);
    }else{
        return nacci(nacci([2,1], n, n), n, len);
    }
}
```
