---
id: 5992e222d397f00d21122931
title: Neno la Fibonacci
challengeType: 1
forumTopicId: 302269
dashedName: fibonacci-word
---

# --description--

Mfululizo wa Maneno ya Fibonacci unaweza kuundwa kwa njia inayofanana na Mfululizo wa Fibonacci, lakini unazingatia kuendelea kuunganisha mfuatano wa herufi.

<pre>Tambua  F_Word<sub>1</sub>  kama  <strong>1</strong>
Tambua  F_Word<sub>2</sub>  kama  <strong>0</strong>
Unda   F_Word<sub>3</sub>  kama  F_Word<sub>2</sub>   kuunganishwa na  F_Word<sub>1</sub>   yaani:  <strong>01</strong>
Unda   F_Word<sub>n</sub>  kama  F_Word<sub>n-1</sub>  kuunganishwa na  F_word<sub>n-2</sub>
</pre>

Hesabu ya entropy inahitajika katika changamoto hii, <a href="https://www.freecodecamp.org/learn/coding-interview-prep/rosetta-code/entropy" target="_blank" rel="noopener noreferrer nofollow">kama inavyoonyeshwa katika changamoto hii ya Rosetta Code</a>

# --instructions--

Andika kitendakazi kurudisha Maneno ya Fibonacci wa kwanza `n`. Idadi ya `n` hutolewa kama kigezo kwa kitendakazi. Kitendakazi kinapaswa kurudisha safu ya vitu. Vitu hivyo vinapaswa kuwa katika muundo: `{ N: 1, Length: 1, Entropy: 0, Word: '1' }`. `Entropy` huhesabiwa kwa mfuatano wa herufi `Word` na kuzungushwa hadi tarakimu 8 za desimali kwa usahihi. Kumbuka kuwa nambari za mfululizo huu huanza kwa `1`.

# --hints--

`fibWord` inapaswa kuwa kitendakazi.

```js
assert(typeof fibWord === 'function');
```

`fibWord(5)` inapaswa kurudisha safu ya data.

```js
assert(Array.isArray(fibWord(5)));
```

`fibWord(5)` inapaswa kurudisha `[{ N:1, Length:1, Entropy:0, Word:"1" },{ N:2, Length:1, Entropy:0, Word:"0" },{ N:3, Length:2, Entropy:1, Word:"01" },{ N:4, Length:3, Entropy:0.91829583, Word:"010" },{ N:5, Length:5, Entropy:0.97095059, Word:"01001" }]`.

```js
assert.deepEqual(fibWord(5), globalThis.words5);
```

`fibWord(7)` inapaswa kurudisha `[{ N:1, Length:1, Entropy:0, Word:"1" },{ N:2, Length:1, Entropy:0, Word:"0" },{ N:3, Length:2, Entropy:1, Word:"01" },{ N:4, Length:3, Entropy:0.91829583, Word:"010" },{ N:5, Length:5, Entropy:0.97095059, Word:"01001" }, { N:6, Length:8, Entropy:0.954434, Word:'01001010' }, { N:7, Length:13, Entropy:0.9612366, Word:'0100101001001' }]`.

```js
assert.deepEqual(fibWord(7), globalThis.words7);
```

# --before-all--

```js
globalThis.words5 = [
  { N: 1, Length: 1, Entropy: 0, Word: '1' },
  { N: 2, Length: 1, Entropy: 0, Word: '0' },
  { N: 3, Length: 2, Entropy: 1, Word: '01' },
  { N: 4, Length: 3, Entropy: 0.91829583, Word: '010' },
  { N: 5, Length: 5, Entropy: 0.97095059, Word: '01001' }
];

globalThis.words7 = [
  { N: 1, Length: 1, Entropy: 0, Word: '1' },
  { N: 2, Length: 1, Entropy: 0, Word: '0' },
  { N: 3, Length: 2, Entropy: 1, Word: '01' },
  { N: 4, Length: 3, Entropy: 0.91829583, Word: '010' },
  { N: 5, Length: 5, Entropy: 0.97095059, Word: '01001' },
  { N: 6, Length: 8, Entropy: 0.954434, Word: '01001010' },
  { N: 7, Length: 13, Entropy: 0.9612366, Word: '0100101001001' }
];
```

# --seed--

## --seed-contents--

```js
function fibWord(n) {

}
```

# --solutions--

```js
// Round to digits
function roundFloat(num, digits) {
  return Math.round(num * 10.0**digits) / (10.0**digits);
}

// Entropy calculation for string with only 0 and 1
function entropy(word) {
  function digitEntropy(count) {
    return count < 1 ? 0
      : - count / word.length * Math.log2(count / word.length);
  }
  const numZeros = word.split('').filter(e => e === '0').length;
  const numOnes  = word.length - numZeros;
  return roundFloat(digitEntropy(numZeros) + digitEntropy(numOnes), 8);
}

// Compute array of Fibonacci words
function fibWord(n) {
  return [...Array(n).keys()]
    .reduce((words, i) => {
      const word = i === 0 ? "1"
                 : i === 1 ? "0"
                 : words[i - 1].Word + words[i - 2].Word;
      words.push(
        { N: i + 1, Length: word.length, Entropy: entropy(word), Word: word }
      );
      return words;
    }, []);
}
```
