---
id: 599d15309e88c813a40baf58
title: Entropy
challengeType: 1
forumTopicId: 302254
dashedName: entropy
---

# --description--

Hesabu entropy ya Shannon H ya mfuatano wa herufi wa ingizo uliotolewa.

Ikiwa na kigezo kisichobadilika cha nasibu $X$ ambacho ni mfuatano wa "alama" $N$ (jumla ya herufi) unaojumuisha herufi tofauti $n$ (n=2 kwa binary), entropy ya Shannon ya X kwa bits/alama ni:

$H_2(X) = -\\sum\_{i=1}^n \\frac{count_i}{N} \\log_2 \\left(\\frac{count_i}{N}\\right)$

ambapo $count_i$ ni idadi ya herufi $n_i$.

# --hints--

`entropy` inapaswa kuwa kitendakazi.

```js
assert(typeof entropy === 'function');
```

`entropy("0")` inapaswa kurudisha `0`

```js
assert.equal(entropy('0'), 0);
```

`entropy("01")` inapaswa kurudisha `1`

```js
assert.equal(entropy('01'), 1);
```

`entropy("0123")` inapaswa kurudisha `2`

```js
assert.equal(entropy('0123'), 2);
```

`entropy("01234567")` inapaswa kurudisha `3`

```js
assert.equal(entropy('01234567'), 3);
```

`entropy("0123456789abcdef")` inapaswa kurudisha `4`

```js
assert.equal(entropy('0123456789abcdef'), 4);
```

`entropy("1223334444")` inapaswa kurudisha `1.8464393446710154`

```js
assert.equal(entropy('1223334444'), 1.8464393446710154);
```

# --seed--

## --seed-contents--

```js
function entropy(s) {

}
```

# --solutions--

```js
function entropy(s) {
    // Create a dictionary of character frequencies and iterate over it.
  function process(s, evaluator) {
    let h = Object.create(null),
      k;
    s.split('').forEach(c => {
      h[c] && h[c]++ || (h[c] = 1); });
    if (evaluator) for (k in h) evaluator(k, h[k]);
    return h;
  }
    // Measure the entropy of a string in bits per symbol.

  let sum = 0,
    len = s.length;
  process(s, (k, f) => {
    const p = f / len;
    sum -= p * Math.log(p) / Math.log(2);
  });
  return sum;
}
```
