---
id: 599d0ba974141b0f508b37d5
title: Nambari msingi emirp
challengeType: 1
forumTopicId: 302253
dashedName: emirp-primes
---

# --description--

Emirp (nambari msingi iliyoandikwa kwa mwelekeo wa nyuma) ni nambari msingi ambazo zinapogeuzwa mwelekeo (katika uwakilishi wao wa desimali) huwa nambari msingi tofauti.

# --instructions--

Andika kitendakazi ambacho:

<ul>
  <li>Inaonyesha nambari emirp za kwanza <code>n</code>.</li>
  <li>Inaonyesha nambari emirp katika safu ya data.</li>
  <li>Inaonyesha idadi ya nambari emirp katika safu ya data.</li>
  <li>Inaonyesha nambari emirp ya <code>n<sup>th</sup></code>.</li>
</ul>

Kitendakazi kinapaswa kupokea vigezo viwili. Kigezo cha kwanza kitapokea `n` au safu ya data kama safu ya data. Kigezo cha pili kitapokea thamani ya kweli au siyo kweli, ambayo inaeleza kama kitendakazi kitarudisha nambari emirp kama safu ya data au nambari moja (idadi ya nambari msingi katika safu ya data au nambari emirp ya <code>n<sup>th</sup></code>). Kulingana na vigezo, kitendakazi kinapaswa kurudisha safu ya data au nambari.

# --hints--

`emirps` inapaswa kuwa kitendakazi.

```js
assert(typeof emirps === 'function');
```

`emirps(20,true)` inapaswa kurudisha `[13,17,31,37,71,73,79,97,107,113,149,157,167,179,199,311,337,347,359,389]`

```js
assert.deepEqual(emirps(20, true), [
  13,
  17,
  31,
  37,
  71,
  73,
  79,
  97,
  107,
  113,
  149,
  157,
  167,
  179,
  199,
  311,
  337,
  347,
  359,
  389
]);
```

`emirps(1000)` inapaswa kurudisha `70529`

```js
assert.deepEqual(emirps(1000), 70529);
```

`emirps([7700,8000],true)` inapaswa kurudisha `[7717,7757,7817,7841,7867,7879,7901,7927,7949,7951,7963]`

```js
assert.deepEqual(emirps([7700, 8000], true), [
  7717,
  7757,
  7817,
  7841,
  7867,
  7879,
  7901,
  7927,
  7949,
  7951,
  7963
]);
```

`emirps([7700,8000],false)` inapaswa kurudisha `11`

```js
assert.deepEqual(emirps([7700, 8000], false), 11);
```

# --seed--

## --seed-contents--

```js
function emirps(n) {

}
```

# --solutions--

```js
function emirps(num, showEmirps)
{
  const is_prime = function(n)
    {
    if (!(n % 2) || !(n % 3)) return false;
    let p = 1;
    while (p * p < n)
                    { if (n % (p += 4) == 0 || n % (p += 2) == 0)
                            { return false; } }
    return true;
  };
  const is_emirp = function(n) {
    const r = parseInt(n.toString().split('').reverse().join(''));
    return r != n && is_prime(n) && is_prime(r);
  };

  let i,
    arr = [];
  if (typeof num === 'number') {
    for (i = 0; arr.length < num; i++) if (is_emirp(i)) arr.push(i);
    // first x emirps
    if (showEmirps) return arr;
    // xth emirp
    return arr.pop();
  }

  if (Array.isArray(num)) {
    for (i = num[0]; i <= num[1]; i++) if (is_emirp(i)) arr.push(i);
    // emirps between x .. y
    if (showEmirps) return arr;
    // number of emirps between x .. y
    return arr.length;
  }
}
```
