---
id: 59c3ec9f15068017c96eb8a3
title: Mfululizo wa Farey
challengeType: 1
forumTopicId: 302266
dashedName: farey-sequence
---

# --description--

Mfululizo wa Farey <code>F<sub>n</sub></code> wa daraja `n` ni mfululizo wa sehemu zilizopunguzwa kabisa kati ya `0` na `1` ambazo, zinapokuwa katika hali ya chini kabisa, zina majina ya chini kuliko au sawa na `n`, zimesanifishwa kwa mpangilio wa ukubwa unaoongezeka.

*Mfululizo wa Farey* mara nyingine huitwa kwa makosa *mfululizo wa Farey*.

Kila mfululizo wa Farey:

<ul>
  <li>huanza na thamani 0, inayowakilishwa na sehemu $ \frac{0}{1} $</li>
  <li>huisha na thamani 1, inayowakilishwa na sehemu $ \frac{1}{1}$.</li>
</ul>

Mifululizo ya Farey ya daraja `1` hadi `5` ni:

<ul>
  <li style='list-style: none;'>${\bf\it{F}}_1 = \frac{0}{1}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_2 = \frac{0}{1}, \frac{1}{2}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_3 = \frac{0}{1}, \frac{1}{3}, \frac{1}{2}, \frac{2}{3}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_4 = \frac{0}{1}, \frac{1}{4}, \frac{1}{3}, \frac{1}{2}, \frac{2}{3}, \frac{3}{4}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_5 = \frac{0}{1}, \frac{1}{5}, \frac{1}{4}, \frac{1}{3}, \frac{2}{5}, \frac{1}{2}, \frac{3}{5}, \frac{2}{3}, \frac{3}{4}, \frac{4}{5}, \frac{1}{1}$</li>
</ul>

# --instructions--

Andika kitendakazi kinachorudisha mfululizo wa Farey wa daraja `n`. Kitendakazi kinapaswa kuwa na kigezo kimoja ambacho ni `n`. Kinapaswa kurudisha mfululizo huo kama safu ya data.

# --hints--

`farey` inapaswa kuwa kitendakazi.

```js
assert(typeof farey === 'function');
```

`farey(3)` inapaswa kurudisha safu ya data

```js
assert(Array.isArray(farey(3)));
```

`farey(3)` inapaswa kurudisha `['0/1','1/3','1/2','2/3','1/1']`

```js
assert.deepEqual(farey(3),['0/1', '1/3', '1/2', '2/3', '1/1']);
```

`farey(4)` inapaswa kurudisha `['0/1','1/4','1/3','1/2','2/3','3/4','1/1']`

```js
assert.deepEqual(farey(4), ['0/1', '1/4', '1/3', '1/2', '2/3', '3/4', '1/1']);
```

`farey(5)` inapaswa kurudisha `['0/1','1/5','1/4','1/3','2/5','1/2','3/5','2/3','3/4','4/5','1/1']`

```js
assert.deepEqual(farey(5), [
  '0/1',
  '1/5',
  '1/4',
  '1/3',
  '2/5',
  '1/2',
  '3/5',
  '2/3',
  '3/4',
  '4/5',
  '1/1'
]);
```

# --seed--

## --seed-contents--

```js
function farey(n) {

}
```

# --solutions--

```js
function farey(n) {
  const sequence = [{ string: "0/1", float: 0.0 }];
  for (let i = 1; i < n; i++) {
    for (let j = n; j >= i; j--) {
      if (i === 1 || j % i > 0) {
        sequence.push({ string: `${i}/${j}`, float: i / j });
      }
    }
  }
  return sequence
    .sort((a, b) => a.float - b.float)
    .map(e => e.string)
}
```
