---
id: 59d9c6bc214c613ba73ff012
title: Nambari za SEDOL
challengeType: 1
forumTopicId: 302305
dashedName: sedols
---

# --description--

<abbr title="Stock Exchange Daily Official List">SEDOL</abbr> ni orodha ya nambari za utambulisho wa dhamana zinazotolewa na Soko la Hisa la London.

Kwa kila nambari ya orodha ya SEDOL za tarakimu 6, hesabu na ongeza nambari ya ukaguzi (checksum). Yaani, ukipokea mfuatano wa herufi wa ingizo upande wa kushoto, kitendakazi chako kinapaswa kurudisha mfuatano unaolingana upande wa kulia:

<pre>
710889 => 7108899
B0YBKJ => B0YBKJ7
406566 => 4065663
B0YBLH => B0YBLH2
228276 => 2282765
B0YBKL => B0YBKL9
557910 => 5579107
B0YBKR => B0YBKR5
585284 => 5852842
B0YBKT => B0YBKT7
B00030 => B000300
</pre>

Kagua kwamba kila ingizo limeundwa kwa usahihi, hasa kuhusiana na herufi halali zinazoruhusiwa katika mfuatano wa SEDOL. Kitendakazi chako kinapaswa kurudisha `null` kwa ingizo lisilo halali.

# --hints--

`sedol` inapaswa kuwa kitendakazi.

```js
assert(typeof sedol === 'function');
```

`sedol('a')` inapaswa kurudisha null.

```js
assert(sedol('a') === null);
```

`sedol('710889')` inapaswa kurudisha '7108899'.

```js
assert(sedol('710889') === '7108899');
```

`sedol('BOATER')` inapaswa kurudisha null.

```js
assert(sedol('BOATER') === null);
```

`sedol('228276')` inapaswa kurudisha '2282765'.

```js
assert(sedol('228276') === '2282765');
```

# --seed--

## --seed-contents--

```js
function sedol(input) {
  const checkSum = 0
  
  return checkSum;
}
```

# --solutions--

```js
function sedol(input) {
  const checkDigit = sedolCheckDigit(input);
  if (checkDigit !== null) {
    return input + checkDigit;
  }
  return null;
}

const weight = [1, 3, 1, 7, 3, 9, 1];
function sedolCheckDigit(char6) {
  if (char6.search(/^[0-9BCDFGHJKLMNPQRSTVWXYZ]{6}$/) === -1) {
    return null;
  }

  let sum = 0;
  for (let i = 0; i < char6.length; i++) {
    sum += weight[i] * parseInt(char6.charAt(i), 36);
  }
  const checkSum = (10 - (sum % 10)) % 10;
  return checkSum.toString();
}
```
