---
id: 5a23c84252665b21eecc7e82
title: Mgawanyiko mkubwa wa pamoja
challengeType: 1
forumTopicId: 302277
dashedName: greatest-common-divisor
---

# --description--

Andika kitendakazi kinachorudisha mgawanyiko mkubwa wa pamoja wa nambari kamili mbili.

# --hints--

`gcd` inapaswa kuwa kitendakazi.

```js
assert(typeof gcd == 'function');
```

`gcd(24,36)` inapaswa kurudisha nambari.

```js
assert(typeof gcd(24, 36) == 'number');
```

`gcd(24,36)` inapaswa kurudisha `12`.

```js
assert.equal(gcd(24, 36), 12);
```

`gcd(30,48)` inapaswa kurudisha `6`.

```js
assert.equal(gcd(30, 48), 6);
```

`gcd(10,15)` inapaswa kurudisha `5`.

```js
assert.equal(gcd(10, 15), 5);
```

`gcd(100,25)` inapaswa kurudisha `25`.

```js
assert.equal(gcd(100, 25), 25);
```

`gcd(13,250)` inapaswa kurudisha `1`.

```js
assert.equal(gcd(13, 250), 1);
```

`gcd(1300,250)` inapaswa kurudisha `50`.

```js
assert.equal(gcd(1300, 250), 50);
```

# --seed--

## --seed-contents--

```js
function gcd(a, b) {

}
```

# --solutions--

```js
function gcd(a, b) {
  return b==0 ? Math.abs(a):gcd(b, a % b);
}
```
