---
id: 5a23c84252665b21eecc7ede
title: Mwaka wa kuruka
challengeType: 1
forumTopicId: 302300
dashedName: leap-year
---

# --description--

Amua kama mwaka uliopewa ni mwaka wa kuruka katika kalenda ya Gregorian.

# --hints--

`isLeapYear` inapaswa kuwa kitendakazi.

```js
assert(typeof isLeapYear == 'function');
```

`isLeapYear()` inapaswa kurudisha thamani ya kweli au siyo kweli.

```js
assert(typeof isLeapYear(2018) == 'boolean');
```

`isLeapYear(2018)` inapaswa kurudisha `false`.

```js
assert.equal(isLeapYear(2018), false);
```

`isLeapYear(2016)` inapaswa kurudisha `true`.

```js
assert.equal(isLeapYear(2016), true);
```

`isLeapYear(2000)` inapaswa kurudisha `true`.

```js
assert.equal(isLeapYear(2000), true);
```

`isLeapYear(1900)` inapaswa kurudisha `false`.

```js
assert.equal(isLeapYear(1900), false);
```

`isLeapYear(1996)` inapaswa kurudisha `true`.

```js
assert.equal(isLeapYear(1996), true);
```

`isLeapYear(1800)` inapaswa kurudisha `false`.

```js
assert.equal(isLeapYear(1800), false);
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

}
```

# --solutions--

```js
function isLeapYear(year) {
  return year % 100 === 0 ? year % 400 === 0 : year % 4 === 0;
}
```
