---
id: 5a23c84252665b21eecc800b
title: Algoriti za kupanga/safu ya pancake
challengeType: 1
forumTopicId: 302315
dashedName: sorting-algorithmspancake-sort
---

# --description--

Andika kitendakazi cha kupanga safu ya data ya nambari kamili (ya ukubwa wowote unaofaa) kwa mpangilio wa kuongezeka kwa kutumia kupanga kwa Pancake. Kitendakazi kinapaswa kurudisha safu ya data iliyopangwa.

Kwa kifupi, badala ya vipengele binafsi kupangwa, operesheni pekee inayoruhusiwa ni "kuzungusha" mwisho mmoja wa orodha, kama ifuatavyo:

<pre>Kabla:
<b>6 7 8 9</b> 2 5 3 4 1<br>
Baada:
<b>9 8 7 6</b> 2 5 3 4 1
</pre>

Mwishoni mwa orodha mmoja tu ndio unaweza kuzungushwa; huu unapaswa kuwa mwisho wa chini, lakini mwisho wa juu pia uko sawa ikiwa ni rahisi zaidi kuandika msimbo au unafanya kazi vizuri zaidi, lakini lazima **iwe** mwisho ule ule kwa suluhisho lote. (Mwishoni mwa orodha unaozungushwa hauwezi kubadilishwa kiholela.)

# --hints--

`pancakeSort` inapaswa kuwa kitendakazi.

```js
assert(typeof pancakeSort == 'function');
```

`pancakeSort([25, 32, 12, 7, 20])` inapaswa kurudisha safu ya data.

```js
assert(Array.isArray(pancakeSort([25, 32, 12, 7, 20])));
```

`pancakeSort([25, 32, 12, 7, 20])` inapaswa kurudisha `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(pancakeSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`pancakeSort([38, 45, 35, 8, 13])` inapaswa kurudisha `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(pancakeSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`pancakeSort([43, 36, 20, 34, 24])` inapaswa kurudisha `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(pancakeSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`pancakeSort([12, 33, 26, 18, 1, 16, 38])` inapaswa kurudisha `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(pancakeSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`pancakeSort([3, 39, 48, 16, 1, 4, 29])` inapaswa kurudisha `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(pancakeSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function pancakeSort(arr) {

}
```

# --solutions--

```js
function pancakeSort(arr) {
  for (var i = arr.length - 1; i >= 1; i--) {
    // find the index of the largest element not yet sorted
    var max_idx = 0;
    var max = arr[0];
    for (var j = 1; j <= i; j++) {
      if (arr[j] > max) {
        max = arr[j];
        max_idx = j;
      }
    }

    if (max_idx == i) continue; // element already in place

    var new_slice;

    // flip arr max element to index 0
    if (max_idx > 0) {
      new_slice = arr.slice(0, max_idx + 1).reverse();
      for (var j = 0; j <= max_idx; j++) arr[j] = new_slice[j];
    }

    // then flip the max element to its place
    new_slice = arr.slice(0, i + 1).reverse();
    for (var j = 0; j <= i; j++) arr[j] = new_slice[j];
  }
  return arr;
}
```
