---
id: 5a23c84252665b21eecc7eb1
title: Matriksi ya utambulisho
challengeType: 1
forumTopicId: 302290
dashedName: identity-matrix
---

# --description--

*Matriksi ya utambulisho* ni matriksi ya mstatili yenye ukubwa wa \\( n \\times n \\), ambapo vipengele vya diagonal vyote ni `1`s (moja), na vipengele vyote vingine ni `0`s (sifuri).

<ul>
  <li style='list-style: none;'>\(\displaystyle I_{n}=\begin{bmatrix} 1 &#x26; 0 &#x26; 0 \cr 0 &#x26; 1 &#x26; 0 \cr 0 &#x26; 0 &#x26; 1 \cr \end{bmatrix}\)</li>
</ul>

# --instructions--

Andika kitendakazi kinachopokea nambari `n` kama kigezo na kurudisha matriksi ya utambulisho ya daraja \\( n \\times n \\).

# --before-each--

```js
let results=[[ [ 1 ] ],
[ [ 1, 0 ], [ 0, 1 ] ],
[ [ 1, 0, 0 ], [ 0, 1, 0 ], [ 0, 0, 1 ] ],
[ [ 1, 0, 0, 0 ], [ 0, 1, 0, 0 ], [ 0, 0, 1, 0 ], [ 0, 0, 0, 1 ] ]]
```

# --hints--

`idMatrix` inapaswa kuwa kitendakazi.

```js
assert(typeof idMatrix == 'function');
```

`idMatrix(1)` inapaswa kurudisha safu ya data.

```js
assert(Array.isArray(idMatrix(1)));
```

`idMatrix(1)` inapaswa kurudisha `[ [ 1 ] ]`.

```js
assert.deepEqual(idMatrix(1), results[0]);
```

`idMatrix(2)` inapaswa kurudisha `[ [ 1, 0 ], [ 0, 1 ] ]`.

```js
assert.deepEqual(idMatrix(2), results[1]);
```

`idMatrix(3)` inapaswa kurudisha `[ [ 1, 0, 0 ], [ 0, 1, 0 ], [ 0, 0, 1 ] ]`.

```js
assert.deepEqual(idMatrix(3), results[2]);
```

`idMatrix(4)` inapaswa kurudisha `[ [ 1, 0, 0, 0 ], [ 0, 1, 0, 0 ], [ 0, 0, 1, 0 ], [ 0, 0, 0, 1 ] ]`.

```js
assert.deepEqual(idMatrix(4), results[3]);
```

# --seed--

## --seed-contents--

```js
function idMatrix(n) {

}
```

# --solutions--

```js
function idMatrix(n) {
    return Array.apply(null, new Array(n)).map(function (x, i, xs) {
        return xs.map(function (_, k) {
            return i === k ? 1 : 0;
        })
    });
}
```
