---
id: 5a23c84252665b21eecc803f
title: Jumlisha tarakimu za nambari kamili
challengeType: 1
forumTopicId: 302331
dashedName: sum-digits-of-an-integer
---

# --description--

Andika kitendakazi kinachopokea mfuatano wa herufi kama kigezo. Mfuatano huu wa herufi unawakilisha nambari inayoweza kuwa katika msingi wowote (chini ya 37) na rudisha jumla ya tarakimu zake.

<ul>
  <li><b>1</b><sub>10</sub> jumlisha kuwa <b>1</b></li>
  <li><b>1234</b><sub>10</sub> jumlisha kuwa <b>10</b></li>
  <li><b>fe</b><sub>16</sub> jumlisha kuwa <b>29</b></li>
  <li><b>f0e</b><sub>16</sub> jumlisha kuwa <b>29</b></li>
</ul>

# --hints--

`sumDigits` inapaswa kuwa kitendakazi.

```js
assert(typeof sumDigits == 'function');
```

`sumDigits("1")` inapaswa kurudisha nambari.

```js
assert(typeof sumDigits('1') == 'number');
```

`sumDigits("1")` inapaswa kurudisha `1`.

```js
assert.equal(sumDigits('1'), 1);
```

`sumDigits("12345")` inapaswa kurudisha `15`.

```js
assert.equal(sumDigits('12345'), 15);
```

`sumDigits("254")` inapaswa kurudisha `11`.

```js
assert.equal(sumDigits('254'), 11);
```

`sumDigits("fe")` inapaswa kurudisha `29`.

```js
assert.equal(sumDigits('fe'), 29);
```

`sumDigits("f0e")` inapaswa kurudisha `29`.

```js
assert.equal(sumDigits('f0e'), 29);
```

`sumDigits("999ABCXYZ")` inapaswa kurudisha `162`.

```js
assert.equal(sumDigits('999ABCXYZ'), 162);
```

# --seed--

## --seed-contents--

```js
function sumDigits(n) {

}
```

# --solutions--

```js
function sumDigits(n) {
  n += '';
  for (var s = 0, i = 0, e = n.length; i < e; i += 1)
    s += parseInt(n.charAt(i), 36);
  return s;
}
```
