---
id: 5a23c84252665b21eecc8013
title: Algoriti za kupanga/sorting strand
challengeType: 1
forumTopicId: 302319
dashedName: sorting-algorithmsstrand-sort
---

# --description--

**Strand sort** huunda seti ndogo zilizopangwa ambazo huunganishwa kuunda matokeo ya mwisho.

Fikiria safu ya data `unsortedArray = [3, 1, 4, 2]`. Chagua kipengee cha kwanza `3` na nakili katika safu ya data tofauti. Tafuta kipengee chochote kikubwa kinachofuata kipengee hiki. Unapopata kipengee kikubwa, katika kesi hii `4`, nakili katika safu ya data tofauti, `[3, 4]`, na linganisha vipengee vinavyofuata na thamani hii mpya, `4`.

Baada ya kufikia mwisho wa safu ya data, toa vipengee ulivyovikili, `[3, 4]`, na anza tena na kipengee cha kwanza kilichobaki katika `unsortedArray`, katika kesi hii `1`.

Kufuata mchakato huu husababisha safu mbili za data zilizopangwa, `[3, 4]` na `[1, 2]`. Unganisha safu hizi mbili kuunda `strandSortedArray`.

```js
const unsortedArray = [3, 1, 4, 2];
const strandsortedArray = [1, 2, 3, 4];
```

Andika kitendakazi cha kupanga safu ya data kwa kutumia **Strand sort**. Kitendakazi kinapaswa kurudisha safu ya data iliyopangwa.


# --hints--

`strandSort` inapaswa kuwa kitendakazi.

```js
assert(typeof strandSort == 'function');
```

`strandSort([25, 32, 12, 7, 20])` inapaswa kurudisha safu ya data.

```js
assert(Array.isArray(strandSort([25, 32, 12, 7, 20])));
```

`strandSort([25, 32, 12, 7, 20])` inapaswa kurudisha `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(strandSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`strandSort([38, 45, 35, 8, 13])` inapaswa kurudisha `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(strandSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`strandSort([43, 36, 20, 34, 24])` inapaswa kurudisha `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(strandSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`strandSort([12, 33, 26, 18, 1, 16, 38])` inapaswa kurudisha `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(strandSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`strandSort([3, 39, 48, 16, 1, 4, 29])` inapaswa kurudisha `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(strandSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function strandSort(list) {

}
```

# --solutions--

```js
function strandSort(list) {
  function merge(left, right) {
    var result = [];
    while (left.length != 0 && right.length != 0) {
      if (left[0] <= right[0]) result.push(left.shift());
      else result.push(right.shift());
    }
    result.push.apply(result, left);
    result.push.apply(result, right);
    return result;
  }

  if (list.length <= 1) return list;
  var result = [];
  while (list.length > 0) {
    var sorted = [];
    sorted.push(list.shift());
    var len = list.length;
    for (var i = 1; i < len; i++) {
      var elem = list[i];
      if (sorted[i - 1] <= elem) {
        sorted.push(elem);
        sorted.splice(i, 1);
      }
    }

    result = merge(result, sorted);
  }
  return result;
}
```
