---
id: 5a23c84252665b21eecc8002
title: Algoriti za kupanga/sorting bogosort
challengeType: 1
forumTopicId: 302311
dashedName: sorting-algorithmsbogosort
---

# --description--

Panga kwa kutumia bogosort safu ya nambari.

Bogosort ni njia ya kuchanganya mkusanyiko kwa bahati nasibu hadi ipangwe.

"Bogosort" ni algoriti isiyofaa kabisa inayotumiwa kama mzaha tu.

Muda wake wa wastani wa utekelezaji ni O(n!) kwa sababu nafasi ya mchanganyiko wowote wa seti kuisha katika mpangilio sahihi ni takriban moja kati ya faktoria ya *n*, na hali mbaya kabisa ni isiyo na kikomo kwa kuwa hakuna uhakika kwamba mchanganyiko wa bahati nasibu utazalisha mfululizo uliopangwa.

Hali yake bora ni O(n) kwa kuwa mzunguko mmoja tu kupitia vipengele unaweza kutosha kuviweka katika mpangilio.

Msimbo wa bandia:

<pre><b>wakati si</b> InOrder(list) <b>fanya</b>
  Shuffle(list)
<b>maliza</b>
</pre>

# --hints--

`bogosort` inapaswa kuwa kitendakazi.

```js
assert(typeof bogosort == 'function');
```

`bogosort([25, 32, 12, 7, 20])` inapaswa kurudisha safu ya data.

```js
assert(Array.isArray(bogosort([25, 32, 12, 7, 20])));
```

`bogosort([25, 32, 12, 7, 20])` inapaswa kurudisha `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(bogosort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`bogosort([38, 45, 35, 8, 13])` inapaswa kurudisha `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(bogosort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`bogosort([43, 36, 20, 34, 24])` inapaswa kurudisha `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(bogosort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`bogosort([12, 33, 26, 18, 1, 16, 38])` inapaswa kurudisha `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(bogosort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`bogosort([3, 39, 48, 16, 1, 4, 29])` inapaswa kurudisha `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(bogosort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function bogosort(v) {

}
```

# --solutions--

```js
function bogosort(v) {
  function shuffle(v) {
    for (
      var j, x, i = v.length;
      i;
      j = Math.floor(Math.random() * i), x = v[--i], v[i] = v[j], v[j] = x
    );
    return v;
  }

  function isSorted(v) {
    for (var i = 1; i < v.length; i++) {
      if (v[i - 1] > v[i]) {
        return false;
      }
    }
    return true;
  }
  var sorted = false;
  while (sorted == false) {
    v = shuffle(v);
    sorted = isSorted(v);
  }
  return v;
}
```
