---
id: 5a23c84252665b21eecc7e78
title: FizzBuzz kwa ujumla
challengeType: 1
forumTopicId: 302273
dashedName: general-fizzbuzz
---

# --description--

Andika toleo la jumla la <a href="https://rosettacode.org/wiki/FizzBuzz" target="_blank" rel="noopener noreferrer nofollow">FizzBuzz</a> ambalo linafanya kazi kwa orodha yoyote ya vigezo, pamoja na maneno yao.

Hii ni utekelezaji wa "fizzbuzz" ambapo sheria za mchezo zinatolewa kwa mtumizi. Unda kitendakazi kutekeleza hili. Kitendakazi kinapaswa kuchukua vigezo viwili.

Kigezo cha kwanza kitakuwa safu ya data yenye sheria za FizzBuzz. Kwa mfano: `[ [3, "Fizz"] , [5, "Buzz"] ]`.

Hii inaonyesha kuwa `Fizz` inapaswa kuchapishwa ikiwa nambari ni mara nyingi ya 3 na `Buzz` ikiwa ni mara nyingi ya 5. Ikiwa ni mara nyingi ya zote mbili basi mfuatano wa herufi unapaswa kuunganishwa kwa mpangilio ulioainishwa katika safu ya data. Katika kesi hii, `FizzBuzz` ikiwa nambari ni mara nyingi ya 3 na 5.

Kigezo cha pili ni nambari ambayo kitendakazi kinapaswa kurudisha mfuatano wa herufi kama ilivyoelezwa hapo juu.

# --hints--

`genFizzBuzz` inapaswa kuwa kitendakazi.

```js
assert(typeof genFizzBuzz == 'function');
```

`genFizzBuzz([[3, "Fizz"],[5, "Buzz"]], 6)` inapaswa kurudisha mfuatano wa herufi.

```js
assert(
  typeof genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz']
    ],
    6
  ) == 'string'
);
```

`genFizzBuzz([[3, "Fizz"],[5, "Buzz"]], 6)` inapaswa kurudisha `"Fizz"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz']
    ],
    6
  ),
  'Fizz'
);
```

`genFizzBuzz([[3, "Fizz"],[5, "Buzz"]], 10)` inapaswa kurudisha `"Buzz"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz']
    ],
    10
  ),
  'Buzz'
);
```

`genFizzBuzz([[3, "Buzz"],[5, "Fizz"]], 12)` inapaswa kurudisha `"Buzz"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Buzz'],
      [5, 'Fizz']
    ],
    12
  ),
  'Buzz'
);
```

`genFizzBuzz([[3, "Buzz"],[5, "Fizz"]], 13)` inapaswa kurudisha `"13"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Buzz'],
      [5, 'Fizz']
    ],
    13
  ),
  '13'
);
```

`genFizzBuzz([[3, "Buzz"],[5, "Fizz"]], 15)` inapaswa kurudisha `"BuzzFizz"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Buzz'],
      [5, 'Fizz']
    ],
    15
  ),
  'BuzzFizz'
);
```

`genFizzBuzz([[3, "Fizz"],[5, "Buzz"]], 15)` inapaswa kurudisha `"FizzBuzz"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz']
    ],
    15
  ),
  'FizzBuzz'
);
```

`genFizzBuzz([[3, "Fizz"],[5, "Buzz"],[7, "Baxx"]], 105)` inapaswa kurudisha `"FizzBuzzBaxx"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz'],
      [7, 'Baxx']
    ],
    105
  ),
  'FizzBuzzBaxx'
);
```

# --seed--

## --seed-contents--

```js
function genFizzBuzz(rules, num) {

}
```

# --solutions--

```js
function genFizzBuzz(rules, num) {
  let res='';
  rules.forEach(function (e) {
    if(num % e[0] == 0)
      res+=e[1];
  })

  if(res==''){
    res=num.toString();
  }

  return res;
}
```
