---
id: 5a23c84252665b21eecc8045
title: Kukata mstatili wa Sutherland-Hodgman
challengeType: 1
forumTopicId: 302336
dashedName: sutherland-hodgman-polygon-clipping
---

# --description--

Algoriti ya kukata ya Sutherland-Hodgman hupata mstatili ambao ni mkusanyiko wa mstatili wowote (mstatili "muhusika") na mstatili wa konvexi (mstatili wa "kukata"). Inatumika katika michoro ya kompyuta (hasa michoro ya 2D) kupunguza ugumu wa mandhari inayotolewa kwa kuondoa sehemu za mstatili ambazo hazihitaji kuonyeshwa. Chukua mstatili uliofungwa unaoelezewa na pointi:

<pre>[(50, 150), (200, 50), (350, 150), (350, 300), (250, 300), (200, 250), (150, 350), (100, 250), (100, 200)]</pre>

na ukate kwa mstatili wa mstatili ulioelezewa na pointi:

<pre>[(100, 100), (300, 100), (300, 300), (100, 300)]</pre>

# --instructions--

Andika kitendakazi kinachopokea safu mbili za data kama vigezo. Safu ya kwanza ina pointi za mstatili wa mhusika na safu ya pili ina pointi za mstatili wa kukata. Kitendakazi kinapaswa kurudisha safu ya data yenye pointi za mstatili uliokatwa. Kila nambari inapaswa kuzungushwa hadi sehemu tatu za desimali.

# --before-each--

```js
const subjectPolygon1 = [
  [50, 150],
  [200, 50],
  [350, 150],
  [350, 300],
  [250, 300],
  [200, 250],
  [150, 350],
  [100, 250],
  [100, 200]
];
const clipPolygon1 = [
  [100, 100],
  [300, 100],
  [300, 300],
  [100, 300]
];
const expected1 = [
  [100, 116.667],
  [125, 100],
  [275, 100],
  [300, 116.667],
  [300, 300],
  [250, 300],
  [200, 250],
  [175, 300],
  [125, 300],
  [100, 250]
];

const subjectPolygon2 = [
  [150, 200],
  [400, 450],
  [30, 50]
];
const clipPolygon2 = [
  [10, 10],
  [300, 200],
  [400, 600],
  [100, 300]
];
const expected2 = [
  [150, 200],
  [350, 400],
  [348.611, 394.444],
  [30, 50]
];

const subjectPolygon3 = [
  [250, 200],
  [100, 450],
  [130, 250]
];
const clipPolygon3 = [
  [50, 60],
  [100, 230],
  [400, 600],
  [100, 300]
];
const expected3 = [
  [129.167, 329.167],
  [119.565, 319.565],
  [121.854, 304.305]
];
```

# --hints--

`clip` inapaswa kuwa kitendakazi.

```js
assert(typeof clip == 'function');
```

`clip([[50, 150], [200, 50], [350, 150], [350, 300], [250, 300], [200, 250], [150, 350], [100, 250], [100, 200]], [[100, 100], [300, 100], [300, 300], [100, 300]])` inapaswa kurudisha safu ya data.

```js
assert(Array.isArray(clip(subjectPolygon1, clipPolygon1)));
```

`clip([[50, 150], [200, 50], [350, 150], [350, 300], [250, 300], [200, 250], [150, 350], [100, 250], [100, 200]], [[100, 100], [300, 100], [300, 300], [100, 300]])` inapaswa kurudisha `[[100, 116.667], [125, 100], [275, 100], [300, 116.667], [300, 300], [250, 300], [200, 250], [175, 300], [125, 300], [100, 250]]`.

```js
assert.deepEqual(clip(subjectPolygon1, clipPolygon1), expected1);
```

`clip([[150, 200], [400, 450], [30, 50]], [[10, 10], [300, 200], [400, 600], [100, 300]])` inapaswa kurudisha `[[150, 200], [350, 400], [348.611, 394.444], [30, 50]]`.

```js
assert.deepEqual(clip(subjectPolygon2, clipPolygon2), expected2);
```

`clip([[250, 200], [100, 450], [130, 250]], [[50, 60], [100, 230], [400, 600], [100, 300]])` inapaswa kurudisha `[[129.167, 329.167], [119.565, 319.565], [121.854, 304.305]]`.

```js
assert.deepEqual(clip(subjectPolygon3, clipPolygon3), expected3);
```

# --seed--

## --seed-contents--

```js
function clip(subjectPolygon, clipPolygon) {

}
```

# --solutions--

```js
function clip(subjectPolygon, clipPolygon) {
  var cp1, cp2, s, e, i, j;
  var inside = function(p) {
    return (
      (cp2[0] - cp1[0]) * (p[1] - cp1[1]) > (cp2[1] - cp1[1]) * (p[0] - cp1[0])
    );
  };
  var intersection = function() {
    var dc = [cp1[0] - cp2[0], cp1[1] - cp2[1]],
      dp = [s[0] - e[0], s[1] - e[1]],
      n1 = cp1[0] * cp2[1] - cp1[1] * cp2[0],
      n2 = s[0] * e[1] - s[1] * e[0],
      n3 = 1.0 / (dc[0] * dp[1] - dc[1] * dp[0]);
    return [(n1 * dp[0] - n2 * dc[0]) * n3, (n1 * dp[1] - n2 * dc[1]) * n3];
  };
  var outputList = subjectPolygon;
  cp1 = clipPolygon[clipPolygon.length - 1];
  for (j in clipPolygon) {
    var cp2 = clipPolygon[j];
    var inputList = outputList;
    outputList = [];
    s = inputList[inputList.length - 1]; //last on the input list
    for (i in inputList) {
      var e = inputList[i];
      if (inside(e)) {
        if (!inside(s)) {
          outputList.push(intersection());
        }
        outputList.push(e);
      } else if (inside(s)) {
        outputList.push(intersection());
      }
      s = e;
    }
    cp1 = cp2;
  }
  return outputList.map(e => e.map(f => Math.round(f * 1000) / 1000));
}
```
