---
id: 5a23c84252665b21eecc7ec2
title: umbali wa Jaro
challengeType: 1
forumTopicId: 302292
dashedName: jaro-distance
---

# --description--

Umbali wa Jaro ni kipimo cha kufanana kati ya mfuatano wa herufi mbili. Kadiri umbali wa Jaro kwa mfuatano wa herufi mbili unavyoongezeka, ndivyo mfuatano huo unavyofanana zaidi. Alama hiyo ni kawaishwa kwa njia kwamba `0` inalingana na kutofanana kabisa na `1` ni mlingano kamili.

**Ufafanuzi**

Umbali wa Jaro \\( d_j \\) wa mfuatano wa herufi mbili \\(s_1\\) na \\(s_2\\) ni

\\begin{align}d_j = \\begin{cases}0& & \\text{ikiwa }m=0 \\\\\\\\{\\frac {1}{3}}\\left({\\frac {m}{|s\_{1}|}}+{\\frac {m}{|s\_{2}|}}+{\\frac {m-t}{m}}\\right)& & \\text{vinginevyo}\\end{cases}\\end{align}

Ambapo:

<ul>
  <li>\(m\) ni idadi ya <i>herufi zinazofanana</i>;</li>
  <li> \(t\) ni nusu ya idadi ya <i>mabadiliko ya nafasi</i>.</li>
</ul>

Herufi mbili kutoka \\(s_1\\) na \\(s_2\\) mtawalia, zinachukuliwa kuwa *zinafanana* tu ikiwa ni sawa na haziko mbali zaidi ya \\(\\left\\lfloor\\frac{\\max(|s_1|,|s_2|)}{2}\\right\\rfloor-1\\).

Kila herufi ya \\(s_1\\) inalinganishwa na herufi zote zinazofanana katika \\(s_2\\). Idadi ya herufi zinazofanana (lakini zikiwa katika mfululizo tofauti) kugawanywa kwa 2 inaelezea idadi ya *mabadiliko ya nafasi*.

**Mfano**

Ikiwa mfuatano wa herufi ni \\(s_1\\) *DWAYNE* na \\(s_2\\) *DUANE* tunapata:

<ul>
  <li>\(m = 4\)</li>
  <li>\(|s_1| = 6\)</li>
  <li>\(|s_2| = 5\)</li>
  <li>\(t = 0\)</li>
</ul>

Tunapata alama ya Jaro: \\(d_j = \\frac{1}{3}\\left(\\frac{4}{6} + \\frac{4}{5} + \\frac{4-0}{4}\\right) = 0.822\\).

# --instructions--

Andika kitendakazi a kinachopokea mfuatano wa herufi mbili kama vigezo na kurudisha umbali wa Jaro unaohusiana.

# --hints--

`jaro` inapaswa kuwa kitendakazi.

```js
assert(typeof jaro == 'function');
```

`jaro("MARTHA", "MARHTA")` inapaswa kurudisha nambari.

```js
assert(typeof jaro('MARTHA', 'MARHTA') == 'number');
```

`jaro("MARTHA", "MARHTA")` inapaswa kurudisha `0.9444444444444445`.

```js
assert.equal(jaro('MARTHA', 'MARHTA'), 0.9444444444444445);
```

`jaro("DIXON", "DICKSONX")` inapaswa kurudisha `0.7666666666666666`.

```js
assert.equal(jaro('DIXON', 'DICKSONX'), 0.7666666666666666);
```

`jaro("JELLYFISH", "SMELLYFISH")` inapaswa kurudisha `0.8962962962962964`.

```js
assert.equal(jaro('JELLYFISH', 'SMELLYFISH'), 0.8962962962962964);
```

`jaro("HELLOS", "CHELLO")` inapaswa kurudisha `0.888888888888889`.

```js
assert.equal(jaro('HELLOS', 'CHELLO'), 0.888888888888889);
```

`jaro("ABCD", "BCDA")` inapaswa kurudisha `0.8333333333333334`.

```js
assert.equal(jaro('ABCD', 'BCDA'), 0.8333333333333334);
```

# --seed--

## --seed-contents--

```js
function jaro(s, t) {

}
```

# --solutions--

```js
function jaro(s, t) {
  var s_len = s.length;
  var t_len = t.length;

  if (s_len == 0 && t_len == 0) return 1;

  var match_distance = Math.max(s_len, t_len) / 2 - 1;

  var s_matches = new Array(s_len);
  var t_matches = new Array(t_len);

  var matches = 0;
  var transpositions = 0;

  for (var i = 0; i < s_len; i++) {
    var start = Math.max(0, i - match_distance);
    var end = Math.min(i + match_distance + 1, t_len);

    for (var j = start; j < end; j++) {
      if (t_matches[j]) continue;
      if (s.charAt(i) != t.charAt(j)) continue;
      s_matches[i] = true;
      t_matches[j] = true;
      matches++;
      break;
    }
  }

  if (matches == 0) return 0;

  var k = 0;
  for (var i = 0; i < s_len; i++) {
    if (!s_matches[i]) continue;
    while (!t_matches[k]) k++;
    if (s.charAt(i) != t.charAt(k)) transpositions++;
    k++;
  }

  return ((matches / s_len) +
    (matches / t_len) +
    ((matches - transpositions / 2.0) / matches)) / 3.0;
}
```
