---
id: 5a23c84252665b21eecc8041
title: Jumla ya mfululizo
challengeType: 1
forumTopicId: 302333
dashedName: sum-of-a-series
---

# --description--

Hesabu nambari ya **n**<sup>th</sup> ya <em>mfululizo</em>, yaani jumla ya nambari **n** za kwanza za <em>mfululizo</em> unaolingana. Kwa kawaida thamani hii, au kikomo chake wakati **n** inakaribia ukomo, pia huitwa *jumla ya mfululizo*, hivyo jina la zoezi hili. Kwa zoezi hili, tumia: $S_n = \displaystyle\sum_{k=1}^n \frac{1}{k^2}$.

# --instructions--

Andika kitendakazi kinachopokea $a$ na $b$ kama vigezo na kurudisha jumla ya wanachama wa mfululizo kuanzia wa $a^{th}$ hadi wa $b^{th}$.

# --hints--

`sum` inapaswa kuwa kitendakazi.

```js
assert(typeof sum == 'function');
```

`sum(1, 100)` inapaswa kurudisha nambari.

```js
assert(typeof sum(1, 100) == 'number');
```

`sum(1, 100)` inapaswa kurudisha `1.6349839001848923`.

```js
assert.equal(sum(1, 100), 1.6349839001848923);
```

`sum(33, 46)` inapaswa kurudisha `0.009262256361481223`.

```js
assert.equal(sum(33, 46), 0.009262256361481223);
```

`sum(21, 213)` inapaswa kurudisha `0.044086990748706555`.

```js
assert.equal(sum(21, 213), 0.044086990748706555);
```

`sum(11, 111)` inapaswa kurudisha `0.08619778593108679`.

```js
assert.equal(sum(11, 111), 0.08619778593108679);
```

`sum(1, 10)` inapaswa kurudisha `1.5497677311665408`.

```js
assert.equal(sum(1, 10), 1.5497677311665408);
```

# --seed--

## --seed-contents--

```js
function sum(a, b) {

}
```

# --solutions--

```js
function sum(a, b) {
  function fn(x) {
    return 1 / (x * x);
  }
  var s = 0;
  for (; a <= b; a++) s += fn(a);
  return s;
}
```
