---
id: 5a23c84252665b21eecc8046
title: Tofauti ya usawa
challengeType: 1
forumTopicId: 16086
dashedName: symmetric-difference
---

# --description--

Ukizingatia seti mbili *A* na *B*, hesabu $(A \\setminus B) \\cup (B \\setminus A).$ Yaani, orodhesha vitu vilivyomo katika *A* au *B* lakini si vyote viwili. Seti hii huitwa tofauti ya usawa ya *A* na *B*. Kwa maneno mengine: $(A \\cup B) \\setminus (A \\cap B)$ (seti ya vitu vilivyomo angalau katika moja kati ya *A* au *B* ikiondoa seti ya vitu vilivyomo katika zote mbili *A* na *B*).

Mfano:

Kwa seti `A = [1, 2, 3]`, na `B = [1, 3, 4]`, tofauti ya usawa ya *A* na *B* ni `[2, 4]`.

# --instructions--

Andika kitendakazi kinachopokea safu mbili za data kama vigezo na kurudisha tofauti ya usawa. Panga safu ya matokeo kabla ya kurudisha.

# --hints--

`symmetricDifference` inapaswa kuwa kitendakazi.

```js
assert(typeof symmetricDifference == 'function');
```

`symmetricDifference(["John", "Bob", "Mary", "Serena"], ["Jim", "Mary", "John", "Bob"])` inapaswa kurudisha safu ya data.

```js
assert(
  Array.isArray(
    symmetricDifference(
      ['John', 'Bob', 'Mary', 'Serena'],
      ['Jim', 'Mary', 'John', 'Bob']
    )
  )
);
```

`symmetricDifference(["John", "Bob", "Mary", "Serena"], ["Jim", "Mary", "John", "Bob"])` inapaswa kurudisha `["Jim", "Serena"]`.

```js
assert.deepEqual(
  symmetricDifference(
    ['John', 'Bob', 'Mary', 'Serena'],
    ['Jim', 'Mary', 'John', 'Bob']
  ),
  ['Jim', 'Serena']
);
```

`symmetricDifference([1, 2, 3], [3, 4])` inapaswa kurudisha `[1, 2, 4]`.

```js
assert.deepEqual(symmetricDifference([1, 2, 3], [3, 4]), [1, 2, 4]);
```

`symmetricDifference([1, 2, 3, 4, 5], [3, 4, 8, 7])` inapaswa kurudisha `[1, 2, 5, 7, 8]`.

```js
assert.deepEqual(symmetricDifference([1, 2, 3, 4, 5], [3, 4, 8, 7]), [
  1,
  2,
  5,
  7,
  8
]);
```

`symmetricDifference([1, 2, 3, 4, 5, 6, 7, 8], [1, 3, 5, 6, 7, 8, 9])` inapaswa kurudisha `[2, 4, 9]`.

```js
assert.deepEqual(
  symmetricDifference([1, 2, 3, 4, 5, 6, 7, 8], [1, 3, 5, 6, 7, 8, 9]),
  [2, 4, 9]
);
```

`symmetricDifference([1, 2, 4, 7, 9], [2, 3, 7, 8, 9])` inapaswa kurudisha `[1, 3, 4, 8]`.

```js
assert.deepEqual(symmetricDifference([1, 2, 4, 7, 9], [2, 3, 7, 8, 9]), [
  1,
  3,
  4,
  8
]);
```

# --seed--

## --seed-contents--

```js
function symmetricDifference(A, B) {

}
```

# --solutions--

```js
function symmetricDifference(A, B) {
  function relative_complement(A, B) {
    return A.filter(function(elem) {
      return B.indexOf(elem) == -1;
    });
  }

  function unique(ary) {
    var u = ary.concat().sort();
    for (var i = 1; i < u.length; ) {
      if (u[i - 1] === u[i]) u.splice(i, 1);
      else i++;
    }
    return u;
  }

  return unique(
    relative_complement(A, B).concat(relative_complement(B, A))
  ).sort();
}
```
