---
id: 5a23c84252665b21eecc8036
title: Ondoa msimbo wa udhibiti na herufi zilizopanuliwa kutoka kwa mfuatano wa herufi
challengeType: 1
forumTopicId: 302327
dashedName: strip-control-codes-and-extended-characters-from-a-string
---

# --description--

Zoezi ni kuondoa msimbo wa udhibiti na herufi zilizopanuliwa kutoka kwa mfuatano wa herufi. Suluhisho linapaswa kuonyesha jinsi ya kufanikisha kila mojawapo ya matokeo yafuatayo: Mfuatano wa herufi wenye msimbo wa udhibiti na herufi zilizopanuliwa zimeondolewa. Katika ASCII, misimbo ya udhibiti ina nambari za desimali kuanzia 0 hadi 31 na 127. Katika mfumo unaotegemea ASCII, ikiwa misimbo ya udhibiti itaondolewa, mfuatano utakuwa na herufi zote zikiwa ndani ya anuwai ya desimali 32 hadi 126 kwenye jedwali la ASCII. Katika mfumo usiotegemea ASCII, tunachukulia herufi ambazo hazina alama inayolingana kwenye jedwali la ASCII (ndani ya anuwai ya desimali 32 hadi 126) kuwa herufi zilizopanuliwa kwa madhumuni ya zoezi hili.

# --hints--

`strip` inapaswa kuwa kitendakazi.

```js
assert(typeof strip == 'function');
```

`strip("abc")` inapaswa kurudisha mfuatano wa herufi.

```js
assert(typeof strip('abc') == 'string');
```

`strip("\ba\x00b\n\rc\fd\xc3")` inapaswa kurudisha `"abcd"`.

```js
assert.equal(strip('\ba\x00b\n\rc\fd\xc3'), 'abcd');
```

`strip("\u0000\n abc\u00E9def\u007F")` inapaswa kurudisha `" abcdef"`.

```js
assert.equal(strip('\u0000\n abc\u00E9def\u007F'), ' abcdef');
```

`strip("a\n\tb\u2102d\u2147f")` inapaswa kurudisha `"abdf"`.

```js
assert.equal(strip('a\n\tb\u2102d\u2147f'), 'abdf');
```

`strip("Français.")` inapaswa kurudisha `"Franais."`.

```js
assert.equal(strip('Français.'), 'Franais.');
```

`strip("123\tabc\u0007DEF\u007F+-*/€æŧðłþ")` inapaswa kurudisha `"123abcDEF+-*/"`.

```js
assert.equal(strip('123\tabc\u0007DEF\u007F+-*/€æŧðłþ'), '123abcDEF+-*/');
```

# --seed--

## --seed-contents--

```js
function strip(s) {

}
```

# --solutions--

```js
function strip(s) {
  return s
    .split('')
    .filter(function(x) {
      var n = x.charCodeAt(0);

      return 31 < n && 127 > n;
    })
    .join('');
}
```
