---
id: 5a23c84252665b21eecc7e03
title: Mzunguko wa kawaida wa jumla
challengeType: 1
forumTopicId: 302240
dashedName: cumulative-standard-deviation
---

# --description--

Fikiria kuwa idadi yote ya watu wa maslahi ni wanafunzi nane katika darasa fulani. Kwa seti ya nambari zisizo na mwisho, mzunguko wa kawaida wa watu hupatikana kwa kuchukua mzizi wa mraba wa wastani wa tofauti za mraba za thamani zilizokatwa kutoka kwa wastani wao. Alama za darasa la wanafunzi nane (yaani, idadi ya watu wa takwimu) ni hizi thamani nane:

$2, 4, 4, 4, 5, 5, 7, 9$

Pointi hizi nane za data zina wastani (wastani) wa 5:

$$\mu ={\frac {2+4+4+4+5+5+7+9}{8}}={\frac {40}{8}}=5$$

Kwanza, hesabu tofauti za kila pointi ya data kutoka kwa wastani, na piga mraba matokeo ya kila moja:

| Tofauti za kila data | Piga mraba matokeo |
| ---------------------- | --------------------|
| $(2-5)^{2}=(-3)^{2}=9$ | $(5-5)^{2}=0^{2}=0$ |
| $(4-5)^{2}=(-1)^{2}=1$ | $(5-5)^{2}=0^{2}=0$ |
| $(4-5)^{2}=(-1)^{2}=1$ | $(7-5)^{2}=2^{2}=4$ |
| $(4-5)^{2}=(-1)^{2}=1$ | $(9-5)^{2}=4^{2}=16$|

Mabadiliko ni wastani wa hizi thamani:

$$\sigma ^{2}={\frac {9+1+1+1+0+0+4+16}{8}}={\frac {32}{8}}=4$$

na mzunguko wa kawaida wa idadi ya watu ni sawa na mzizi wa mraba wa mabadiliko:

$$\sigma ={\sqrt {4}}=2$$

Andika kitendakazi kinachopokea safu ya nambari kama kigezo na kurudisha <a href="https://rosettacode.org/wiki/Standard_deviation" target="_blank" rel="noopener noreferrer nofollow">mzunguko wa kawaida</a> wa mfululizo huo.

# --hints--

`standardDeviation` inapaswa kuwa kitendakazi.

```js
assert(typeof standardDeviation == 'function');
```

`standardDeviation([2, 4, 4, 4, 5, 5, 7, 9])` inapaswa kurudisha nambari.

```js
assert(typeof standardDeviation([2, 4, 4, 4, 5, 5, 7, 9]) == 'number');
```

`standardDeviation([2, 4, 4, 4, 5, 5, 7, 9])` inapaswa kurudisha `2`.

```js
assert.equal(standardDeviation([2, 4, 4, 4, 5, 5, 7, 9]), 2);
```

`standardDeviation([600, 470, 170, 430, 300])` inapaswa kurudisha `147.323`.

```js
assert.equal(standardDeviation([600, 470, 170, 430, 300]), 147.323);
```

`standardDeviation([75, 83, 96, 100, 121, 125])` inapaswa kurudisha `18.239`.

```js
assert.equal(standardDeviation([75, 83, 96, 100, 121, 125]), 18.239);
```

`standardDeviation([23, 37, 45, 49, 56, 63, 63, 70, 72, 82])` inapaswa kurudisha `16.87`.

```js
assert.equal(
  standardDeviation([23, 37, 45, 49, 56, 63, 63, 70, 72, 82]),
  16.87
);
```

`standardDeviation([271, 354, 296, 301, 333, 326, 285, 298, 327, 316, 287, 314])` inapaswa kurudisha `22.631`.

```js
assert.equal(
  standardDeviation([
    271,
    354,
    296,
    301,
    333,
    326,
    285,
    298,
    327,
    316,
    287,
    314
  ]),
  22.631
);
```

# --seed--

## --seed-contents--

```js
function standardDeviation(arr) {

}
```

# --solutions--

```js
function standardDeviation(arr) {
  var sum = 0,
    sum_sq = 0,
    n = arr.length;
  arr.forEach(function(e) {
    sum += e;
    sum_sq += e * e;
  });

  var std_dev = Math.sqrt(sum_sq / n - Math.pow(sum / n, 2));
  return Math.round(std_dev * 1000) / 1000;
}
```
