---
id: 5a23c84252665b21eecc8005
title: Algoriti za kupanga/safu ya data ya comb
challengeType: 1
forumTopicId: 302313
dashedName: sorting-algorithmscomb-sort
---

# --description--

Tekeleza *safu ya data ya comb*.

**Safu ya data ya Comb** ni aina ya Bubble Sort.

Kama Shell sort, Comb Sort huongeza pengo linalotumika katika kulinganisha na kubadilishana.

Kugawanya pengo kwa $(1-e^{-\\varphi})^{-1} \\approx 1.247330950103979$ ni njia bora zaidi, lakini 1.3 inaweza kuwa rahisi zaidi.

Baadhi ya utekelezaji hutumia insertion sort mara pengo linapokuwa chini ya kiwango fulani.

Aina:

<ul>
  <li>Combsort11 huhakikisha pengo linaishia (11, 8, 6, 4, 3, 2, 1), ambayo ni haraka sana kuliko mwisho mwingine wa uwezekano.</li>
  <li>Combsort yenye mwisho tofauti hubadilika kuwa aina bora zaidi ya kupanga wakati data iko karibu kupanga (pengonipo ndogo). Safu ya data ya comb yenye pengo ndogo si bora sana kuliko Bubble Sort.</li>
</ul>

Msimbo wa bandia:

<pre><b>function</b> combsort(<b>array</b> input)
  gap := input<b>.size</b> <i>//anzisha ukubwa wa pengo</i>
  <b>loop until</b> gap = 1 <b>and</b> swaps = 0
    <i>//sasisha thamani ya pengo kwa comb inayofuata. Hapa ni mfano</i>
    gap := int(gap / 1.25)
    <b>if</b> gap &#x3C; 1 
      <i>//pengoncha ni 1</i>
      gap := 1
    <b>end if</b>
    i := 0
    swaps := 0 <i>//angalia <a href='https://rosettacode.org/wiki/Sorting_algorithms/Bubble_sort' target='_blank'>Bubble Sort</a> kwa maelezo</i>
    <i>//"comb" moja juu ya orodha ya ingizo</i>
    <b>loop until</b> i + gap >= input<b>.size</b> <i>//angalia <a href='https://rosettacode.org/wiki/Sorting_algorithms/Shell_sort' target='_blank'>Shell sort</a> kwa wazo sawa</i>
      <b>if</b> input[i] > input[i+gap]
        <b>swap</b>(input[i], input[i+gap])
        swaps := 1 <i>// Onyesha kuwa kubadilishana kumetokea, hivyo</i>
            <i>// orodha haijahakikishiwa kuwa imepangwa</i>
      <b>end if</b>
      i := i + 1
    <b>end loop</b>
  <b>end loop</b>
<b>end function</b>
</pre>

# --instructions--

Andika kitendakazi kinachopanga safu ya data iliyotolewa kwa kutumia safu ya data ya Comb.

# --hints--

`combSort` inapaswa kuwa kitendakazi.

```js
assert(typeof combSort == 'function');
```

`combSort([25, 32, 12, 7, 20])` inapaswa kurudisha safu ya data.

```js
assert(Array.isArray(combSort([25, 32, 12, 7, 20])));
```

`combSort([25, 32, 12, 7, 20])` inapaswa kurudisha `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(combSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`combSort([38, 45, 35, 8, 13])` inapaswa kurudisha `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(combSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`combSort([43, 36, 20, 34, 24])` inapaswa kurudisha `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(combSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`combSort([12, 33, 26, 18, 1, 16, 38])` inapaswa kurudisha `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(combSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`combSort([3, 39, 48, 16, 1, 4, 29])` inapaswa kurudisha `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(combSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function combSort(arr) {

}
```

# --solutions--

```js
function combSort(arr) {
  function is_array_sorted(arr) {
    var sorted = true;
    for (var i = 0; i < arr.length - 1; i++) {
      if (arr[i] > arr[i + 1]) {
        sorted = false;
        break;
      }
    }
    return sorted;
  }
  var iteration_count = 0;
  var gap = arr.length - 2;
  var decrease_factor = 1.25;

  // Until array is not sorted, repeat iterations
  while (!is_array_sorted(arr)) {
    // If not first gap
    if (iteration_count > 0)
      // Calculate gap
      gap = gap == 1 ? gap : Math.floor(gap / decrease_factor);

    // Set front and back elements and increment to a gap
    var front = 0;
    var back = gap;
    while (back <= arr.length - 1) {
      // If elements are not ordered swap them
      if (arr[front] > arr[back]) {
        var temp = arr[front];
        arr[front] = arr[back];
        arr[back] = temp;
      }

      // Increment and re-run swapping
      front += 1;
      back += 1;
    }
    iteration_count += 1;
  }

  return arr;
}
```
