---
id: 5a23c84252665b21eecc8014
title: Utulivu wa kupanga
challengeType: 1
forumTopicId: 302308
dashedName: sort-stability
---

# --description--

Unapopanga rekodi katika jedwali kwa safu au uwanja fulani, kupanga kwa utulivu (stable sort) kila mara kutahifadhi mpangilio wa rekodi zilizo na ufunguo sawa.

Kwa mfano, katika jedwali hili la nchi na miji, kupanga kwa utulivu kwa safu ya **pili**, miji, kutahifadhi Birmingham ya Marekani juu ya Birmingham ya Uingereza. (Ingawa kupanga kisivyo na utulivu *kinaweza*, katika kesi hii, kuweka Birmingham ya Marekani juu ya Birmingham ya Uingereza, utaratibu wa kupanga kwa utulivu utahakikisha hilo).

<pre>UK  London
US  New York
US  Birmingham
UK  Birmingham
</pre>

Vivyo hivyo, kupanga kwa utulivu kwa safu ya kwanza tu kutazalisha "UK London" kama kipengele cha kwanza na "US Birmingham" kama kipengele cha mwisho (kwa sababu mpangilio wa vipengele vyenye neno la kwanza sawa – "UK" au "US" – utahifadhiwa).

# --instructions--

Andika kitendakazi kinachopokea safu ya data ya vipimo viwili kama kigezo. Kila kipengele kina vipengele viwili kama mfano ulio hapo juu. Kitendakazi kinapaswa kupanga safu ya data kama ilivyoelezwa hapo awali na kurudisha safu iliyopangwa.

# --hints--

`stableSort` inapaswa kuwa kitendakazi.

```js
assert(typeof stableSort == 'function');
```

`stableSort([["UK", "London"], ["US", "New York"], ["US", "Birmingham"], ["UK", "Birmingham"]])` inapaswa kurudisha safu ya data.

```js
assert(
  Array.isArray(
    stableSort([
      ['UK', 'London'],
      ['US', 'New York'],
      ['US', 'Birmingham'],
      ['UK', 'Birmingham']
    ])
  )
);
```

`stableSort([["UK", "London"], ["US", "New York"], ["US", "Birmingham"], ["UK", "Birmingham"]])` inapaswa kurudisha `[["US", "Birmingham"], ["UK", "Birmingham"], ["UK", "London"], ["US", "New York"]]`.

```js
assert.deepEqual(
  stableSort([
    ['UK', 'London'],
    ['US', 'New York'],
    ['US', 'Birmingham'],
    ['UK', 'Birmingham']
  ]),
  [
    ['US', 'Birmingham'],
    ['UK', 'Birmingham'],
    ['UK', 'London'],
    ['US', 'New York']
  ]
);
```

`stableSort([[2, 2], [1, 2], [1, 4], [1, 5]])` inapaswa kurudisha `[[2, 2], [1, 2], [1, 4], [1, 5]]`.

```js
assert.deepEqual(
  stableSort([
    [2, 2],
    [1, 2],
    [1, 4],
    [1, 5]
  ]),
  [
    [2, 2],
    [1, 2],
    [1, 4],
    [1, 5]
  ]
);
```

`stableSort([[11, 55], [12, 45], [11, 45], [32, 45]])` inapaswa kurudisha `[[12, 45], [11, 45], [32, 45], [11, 55]]`.

```js
assert.deepEqual(
  stableSort([
    [11, 55],
    [12, 45],
    [11, 45],
    [32, 45]
  ]),
  [
    [12, 45],
    [11, 45],
    [32, 45],
    [11, 55]
  ]
);
```

`stableSort([[10, 22], [1, 2], [1, 4], [1, 5], [10, 9]])` inapaswa kurudisha `[[1, 2], [1, 4], [1, 5], [10, 9], [10, 22]]`.

```js
assert.deepEqual(
  stableSort([
    [10, 22],
    [1, 2],
    [1, 4],
    [1, 5],
    [10, 9]
  ]),
  [
    [1, 2],
    [1, 4],
    [1, 5],
    [10, 9],
    [10, 22]
  ]
);
```

`stableSort([[55, 54], [12, 22], [31, 43], [31, 54], [10, 49]])` inapaswa kurudisha `[[12, 22], [31, 43], [10, 49], [55, 54], [31, 54]]`.

```js
assert.deepEqual(
  stableSort([
    [55, 54],
    [12, 22],
    [31, 43],
    [31, 54],
    [10, 49]
  ]),
  [
    [12, 22],
    [31, 43],
    [10, 49],
    [55, 54],
    [31, 54]
  ]
);
```

# --seed--

## --seed-contents--

```js
function stableSort(arr) {

}
```

# --solutions--

```js
function stableSort(arr) {
  arr.sort(function(a, b) {
    return a[1] < b[1] ? -1 : a[1] > b[1] ? 1 : 0;
  });
  return arr;
}
```
