---
id: 5a23c84252665b21eecc7ec1
title: Kuweka mraba wa nambari rudufu
challengeType: 1
forumTopicId: 302291
dashedName: iterated-digits-squaring
---

# --description--

Ikiwa utaongeza mraba wa tarakimu za nambari asilia (nambari kamili kubwa kuliko sifuri), daima utamaliza kwa 1 au 89:

<pre>15 -> 26 -> 40 -> 16 -> 37 -> 58 -> 89
7 -> 49 -> 97 -> 130 -> 10 -> 1
</pre>

# --instructions--

Andika kitendakazi kinachopokea nambari kama kigezo na kurudisha 1 au 89 baada ya kufanya mchakato ulioelezwa.

# --hints--

`iteratedSquare` inapaswa kuwa kitendakazi.

```js
assert(typeof iteratedSquare == 'function');
```

`iteratedSquare(4)` inapaswa kurudisha nambari.

```js
assert(typeof iteratedSquare(4) == 'number');
```

`iteratedSquare(4)` inapaswa kurudisha `89`.

```js
assert.equal(iteratedSquare(4), 89);
```

`iteratedSquare(7)` inapaswa kurudisha `1`.

```js
assert.equal(iteratedSquare(7), 1);
```

`iteratedSquare(15)` inapaswa kurudisha `89`.

```js
assert.equal(iteratedSquare(15), 89);
```

`iteratedSquare(20)` inapaswa kurudisha `89`.

```js
assert.equal(iteratedSquare(20), 89);
```

`iteratedSquare(70)` inapaswa kurudisha `1`.

```js
assert.equal(iteratedSquare(70), 1);
```

`iteratedSquare(100)` inapaswa kurudisha `1`.

```js
assert.equal(iteratedSquare(100), 1);
```

# --seed--

## --seed-contents--

```js
function iteratedSquare(n) {

}
```

# --solutions--

```js
function iteratedSquare(n) {
    var total;
    while (n != 89 && n != 1) {
        total = 0;
        while (n > 0) {
            total += Math.pow(n % 10, 2);
            n = Math.floor(n/10);
        }
        n = total;
    }
    return n;
}
```
