---
id: 5a23c84252665b21eecc7eca
title: Nambari za Kaprekar
challengeType: 1
forumTopicId: 302296
dashedName: kaprekar-numbers
---

# --description--

Nambari kamili chanya ni nambari ya Kaprekar ikiwa:

<ul>
  <li>Ni 1, au,</li>
  <li>Uwawakilishi wa desimali wa mraba wake unaweza kugawanywa mara moja katika sehemu mbili zinazojumuisha nambari kamili chanya ambazo jumla yao ni nambari asili. </li>
</ul>

Kumbuka kuwa mgawanyiko unaosababisha sehemu yenye 0 tu haukubaliki, kwa sababu 0 haizingatiwi kama chanya.

Mfano wa nambari za Kaprekar:

<ul>
  <li><code>2223</code> ni nambari ya Kaprekar, kwa sababu <code>2223 * 2223 = 4941729</code>, <code>4941729</code> inaweza kugawanywa kuwa <code>494</code> na <code>1729</code>, na <code>494 + 1729 = 2223</code></li>
  <li>Mfululizo wa nambari za Kaprekar unajulikana kama A006886, na unaanza kama <code>1, 9, 45, 55, ...</code></li>
</ul>

# --instructions--

Andika kitendakazi kinachopokea nambari $n$, msingi $bs$, na kurudisha kweli ikiwa nambari ni nambari ya Kaprekar kwa msingi uliotolewa. Vinginevyo, kitendakazi kirudishe si kweli.

# --hints--

`isKaprekar` inapaswa kuwa kitendakazi.

```js
assert(typeof isKaprekar == 'function');
```

`isKaprekar(1, 10)` inapaswa kurudisha thamani ya kweli au siyo kweli.

```js
assert(typeof isKaprekar(1, 10) == 'boolean');
```

`isKaprekar(1, 10)` inapaswa kurudisha `true`.

```js
assert.equal(isKaprekar(1, 10), true);
```

`isKaprekar(9, 10)` inapaswa kurudisha `true`.

```js
assert.equal(isKaprekar(9, 10), true);
```

`isKaprekar(2223, 10)` inapaswa kurudisha `true`.

```js
assert.equal(isKaprekar(2223, 10), true);
```

`isKaprekar(22823, 10)` inapaswa kurudisha `false`.

```js
assert.equal(isKaprekar(22823, 10), false);
```

`isKaprekar(9, 17)` inapaswa kurudisha `false`.

```js
assert.equal(isKaprekar(9, 17), false);
```

`isKaprekar(225, 17)` inapaswa kurudisha `true`.

```js
assert.equal(isKaprekar(225, 17), true);
```

`isKaprekar(999, 17)` inapaswa kurudisha `false`.

```js
assert.equal(isKaprekar(999, 17), false);
```

# --seed--

## --seed-contents--

```js
function isKaprekar(n, bs) {

}
```

# --solutions--

```js
function isKaprekar(n, bs) {
  if (n < 1) return false;
  if (n == 1) return true;
  for (var a = n * n, b = 0, s = 1; a; s *= bs) {
    b += (a % bs) * s;
    a = Math.floor(a / bs);
    if (b && a + b == n) return true;
  }
  return false;
}
```
