---
id: 5a23c84252665b21eecc8010
title: Algoriti za kupanga/shell sort
challengeType: 1
forumTopicId: 302317
dashedName: sorting-algorithmsshell-sort
---

# --description--

Andika kitendakazi cha kupanga safu ya data ya vipengele kwa kutumia algoriti ya Shell sort, aina ya kupanga kwa kupunguza hatua. Kitendakazi kinapaswa kurudisha safu ya data iliyopangwa.

Shell sort (inayojulikana pia kama Shellsort au njia ya Shell) imepewa jina la mvumbuzi wake, Donald Shell, ambaye alichapisha algoriti hii mwaka 1959.

Shell sort ni mfululizo wa kupanga kwa kuingiza vipande vilivyopangwa kwa mfululizo wa hatua. Ukubwa wa hatua hupunguzwa baada ya kila mzunguko hadi hatua iwe 1.

Kwa hatua ya 1, kupanga ni aina ya kawaida ya kuingiza vipande, lakini wakati huu data inahakikishiwa kuwa karibu imepangwa, ambayo ni "hali nzuri zaidi" ya kuingiza vipande.

Mfululizo wowote utaweza kupanga data mradi tu unamalizika kwa 1, lakini baadhi hufanya kazi vizuri zaidi kuliko mengine.

Utafiti wa majaribio umeonyesha kuwa mfululizo wa hatua wa kijiometri wenye uwiano wa takriban 2.2 hufanya kazi vizuri katika vitendo.

# --hints--

`shellSort` inapaswa kuwa kitendakazi.

```js
assert(typeof shellSort == 'function');
```

`shellSort([25, 32, 12, 7, 20])` inapaswa kurudisha safu ya data.

```js
assert(Array.isArray(shellSort([25, 32, 12, 7, 20])));
```

`shellSort([25, 32, 12, 7, 20])` inapaswa kurudisha `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(shellSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`shellSort([38, 45, 35, 8, 13])` inapaswa kurudisha `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(shellSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`shellSort([43, 36, 20, 34, 24])` inapaswa kurudisha `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(shellSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`shellSort([12, 33, 26, 18, 1, 16, 38])` inapaswa kurudisha `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(shellSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`shellSort([3, 39, 48, 16, 1, 4, 29])` inapaswa kurudisha `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(shellSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function shellSort(a) {

}
```

# --solutions--

```js
function shellSort(a) {
  for (var h = a.length; h > 0; h = parseInt(h / 2)) {
    for (var i = h; i < a.length; i++) {
      var k = a[i];
      for (var j = i; j >= h && k < a[j - h]; j -= h) a[j] = a[j - h];
      a[j] = k;
    }
  }
  return a;
}
```
