---
id: 5a23c84252665b21eecc8017
title: Soundex
challengeType: 1
forumTopicId: 302320
dashedName: soundex
---

# --description--

**Algoriti ya Soundex** inahusu *nia* za maneno. Inaunda uwakilishi wa maneno yanayosikika kwa njia sawa.

Inatumika kwa utatafutaji wa <em>majina</em> na <em>anwani</em>. Hii ina maana mtu aliyekamilisha <em>jina</em> anaweza makini na jinsi linavyosikika badala ya kusahihisha tahajia za <em>majina</em>.

Kwa mfano:

Ikiwa unasikia jina `Quenci` kwa mara ya kwanza, na ukalitahajia vibaya, utapata msimbo wa **Soundex** `Q520`.

Unapotahajia jina `Quincy` kwa usahihi mara nyingine, bado utapata msimbo uleule `Q520`, ambao una maana unaweza kuunganisha matamshi mbalimbali ya jina moja kwa mtu mmoja bila haja ya kuongeza kila tahajia.

Hizi ndizo kanuni:

<ul>
  <li>Kama vokali (A, E, I, O, U) inatenganisha konsonanti mbili zenye msimbo wa soundex sawa, konsonanti iliyo upande wa kulia wa vokali hiyo ndio itapewa msimbo. Tymczak hupatiwa msimbo T-522 (T, 5 kwa M, 2 kwa C, Z haizingatiwi (angalia kanuni ya "Kando kwa Kando" hapo juu), 2 kwa K). Kwa kuwa vokali "A" inatenganisha Z na K, K hupatiwa msimbo.</li>
  <li>Kama "H" au "W" zinatenganisha konsonanti mbili zenye msimbo wa soundex sawa, konsonanti iliyo upande wa kulia wa vokali haipatiwi msimbo. Mfano: Ashcraft hupatiwa msimbo A-261 (A, 2 kwa S, C haizingatiwi, 6 kwa R, 1 kwa F). Haipatiwi msimbo A-226.</li>
</ul>

# --instructions--

Andika kitendakazi kinachopokea mfuatano wa herufi kama kigezo na kurudisha mfuatano wa herufi uliobadilishwa kuwa msimbo.

# --hints--

`soundex` inapaswa kuwa kitendakazi.

```js
assert(typeof soundex == 'function');
```

`soundex("Soundex")` inapaswa kurudisha mfuatano wa herufi.

```js
assert(typeof soundex('Soundex') == 'string');
```

`soundex("Soundex")` inapaswa kurudisha `"S532"`.

```js
assert.equal(soundex('Soundex'), 'S532');
```

`soundex("Example")` inapaswa kurudisha `"E251"`.

```js
assert.equal(soundex('Example'), 'E251');
```

`soundex("Sownteks")` inapaswa kurudisha `"S532"`.

```js
assert.equal(soundex('Sownteks'), 'S532');
```

`soundex("Ekzampul")` inapaswa kurudisha `"E251"`.

```js
assert.equal(soundex('Ekzampul'), 'E251');
```

`soundex("Euler")` inapaswa kurudisha `"E460"`.

```js
assert.equal(soundex('Euler'), 'E460');
```

`soundex("Gauss")` inapaswa kurudisha `"G200"`.

```js
assert.equal(soundex('Gauss'), 'G200');
```

`soundex("Hilbert")` inapaswa kurudisha `"H416"`.

```js
assert.equal(soundex('Hilbert'), 'H416');
```

`soundex("Knuth")` inapaswa kurudisha `"K530"`.

```js
assert.equal(soundex('Knuth'), 'K530');
```

`soundex("Lloyd")` inapaswa kurudisha `"L300"`.

```js
assert.equal(soundex('Lloyd'), 'L300');
```

`soundex("Lukasiewicz")` inapaswa kurudisha `"L222"`.

```js
assert.equal(soundex('Lukasiewicz'), 'L222');
```

# --seed--

## --seed-contents--

```js
function soundex(s) {

}
```

# --solutions--

```js
function soundex(s) {
  var a = s.toLowerCase().split('');
  var f = a.shift(),
    r = '',
    codes = {
      a: '',
      e: '',
      i: '',
      o: '',
      u: '',
      b: 1,
      f: 1,
      p: 1,
      v: 1,
      c: 2,
      g: 2,
      j: 2,
      k: 2,
      q: 2,
      s: 2,
      x: 2,
      z: 2,
      d: 3,
      t: 3,
      l: 4,
      m: 5,
      n: 5,
      r: 6
    };
  r =
    f +
    a
      .map(function(v, i, a) {
        return codes[v];
      })
      .filter(function(v, i, a) {
        return i === 0 ? v !== codes[f] : v !== a[i - 1];
      })
      .join('');

  return (r + '000').slice(0, 4).toUpperCase();
}
```
