---
id: 5a23c84252665b21eecc7eb0
title: I kabla ya E isipokuwa baada ya C
challengeType: 1
forumTopicId: 302288
dashedName: i-before-e-except-after-c
---

# --description--

**<a href="http://www.rosettacode.org/wiki/I_before_E_except_after_C" target="_blank" rel="noopener noreferrer nofollow">"I kabla ya E, isipokuwa baada ya C"</a>** ni kanuni ya jumla ya tahajia ya lugha ya Kiingereza. Ikiwa mtu hana uhakika kama neno linaandikwa na digraph `ei` au `ie`, wimbo huu unapendekeza kuwa mpangilio sahihi ni `ie` isipokuwa herufi inayotangulia ni `c`, ambapo inaweza kuwa `ei`.

Kwa kutumia maneno yaliyotolewa, hakikisha kama vifungu vidogo viwili vya msemo vinaweza kuwa vya kweli peke yao:

<ol>
  <li>
    <i>"I kabla ya E wakati haikutanguliwi na C".</i>
  </li>
  <li>
    <i>"E kabla ya I wakati ikitanguliwa na C".</i>
  </li>
</ol>

Kama vifungu vidogo vyote viwili vinaweza kuwa vya kweli basi msemo wa awali unaweza kusemwa kuwa ni wa kweli.

# --instructions--

Andika kitendakazi kinachokubali neno na hakiki kama neno linafuata kanuni hii. Kitendakazi kinapaswa kurudisha true kama neno linafuata kanuni na false kama halifuati.

# --hints--

`IBeforeExceptC` inapaswa kuwa kitendakazi.

```js
assert(typeof IBeforeExceptC == 'function');
```

`IBeforeExceptC("receive")` inapaswa kurudisha thamani ya kweli au siyo kweli.

```js
assert(typeof IBeforeExceptC('receive') == 'boolean');
```

`IBeforeExceptC("receive")` inapaswa kurudisha `true`.

```js
assert.equal(IBeforeExceptC('receive'), true);
```

`IBeforeExceptC("science")` inapaswa kurudisha `false`.

```js
assert.equal(IBeforeExceptC('science'), false);
```

`IBeforeExceptC("imperceivable")` inapaswa kurudisha `true`.

```js
assert.equal(IBeforeExceptC('imperceivable'), true);
```

`IBeforeExceptC("inconceivable")` inapaswa kurudisha `true`.

```js
assert.equal(IBeforeExceptC('inconceivable'), true);
```

`IBeforeExceptC("insufficient")` inapaswa kurudisha `false`.

```js
assert.equal(IBeforeExceptC('insufficient'), false);
```

`IBeforeExceptC("omniscient")` inapaswa kurudisha `false`.

```js
assert.equal(IBeforeExceptC('omniscient'), false);
```

# --seed--

## --seed-contents--

```js
function IBeforeExceptC(word) {

}
```

# --solutions--

```js
function IBeforeExceptC(word)
{
    if(word.indexOf("c")==-1 && word.indexOf("ie")!=-1)
        return true;
    else if(word.indexOf("cei")!=-1)
        return true;
    return false;
}
```
