---
id: 5a23c84252665b21eecc8001
title: Algoriti za kupanga/safu ya data ya bead sort
challengeType: 1
forumTopicId: 302310
dashedName: sorting-algorithmsbead-sort
---

# --description--

*Bead sort* huanza kwa kuunda matriki ya sifuri zenye urefu sawa na thamani ya kipengele kikubwa zaidi katika safu ya data ya ingizo. Matriki hubadilishwa kwa kuongeza moja kwa vipengele vyote kati ya kielezo cha sifuri na kielezo kinachotajwa na kipengele cha sasa. Mchakato huu hurudiwa, hadi matriki ijazwe.

Kupitia matriki, kujumlisha idadi ya vipengele vinavyozidi sifuri, kisha kupunguza thamani ya kila kipengele kwa moja hutoa safu ya data iliyopangwa.

**Kumbuka:** Kila kipengele katika safu ya data ya ingizo ni cha kipekee.

Panga safu ya data ya nambari kamili chanya kwa kutumia Algoriti ya Bead Sort.

# --hints--

`beadSort` inapaswa kuwa kitendakazi.

```js
assert(typeof beadSort == 'function');
```

`beadSort([25, 32, 12, 7, 20])` inapaswa kurudisha safu ya data.

```js
assert(Array.isArray(beadSort([25, 32, 12, 7, 20])));
```

`beadSort([25, 32, 12, 7, 20])` inapaswa kurudisha `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(beadSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`beadSort([38, 45, 35, 8, 13])` inapaswa kurudisha `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(beadSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`beadSort([43, 36, 20, 34, 24])` inapaswa kurudisha `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(beadSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`beadSort([12, 33, 26, 18, 1, 16, 38])` inapaswa kurudisha `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(beadSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`beadSort([3, 39, 48, 16, 1, 4, 29])` inapaswa kurudisha `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(beadSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function beadSort(arr) {

}
```

# --solutions--

```js
function beadSort(arr) {
  let max = 0;
  for (let i = 0; i < arr.length; i++) if (arr[i] > max) max = arr[i];
  const grid = new Array(arr.length);
  for (let i = 0; i < grid.length; i++) {
    grid[i] = new Array(max);
  }
  const levelcount = new Array(max);
  levelcount.fill(0);
  for (let i = 0; i < max; i++) {
    levelcount[i] = 0;
    for (let j = 0; j < arr.length; j++) grid[j][i] = '_';
  }
  for (let i = 0; i < arr.length; i++) {
    let num = arr[i];
    for (let j = 0; num > 0; j++) {
      grid[levelcount[j]++][j] = '*';
      num--;
    }
  }
  const sorted = new Array(arr.length);
  sorted.fill(0);
  for (let i = 0; i < arr.length; i++) {
    let putt = 0;
    for (
      let j = 0;
      j < max &&
      (function(c) {
        return c.charCodeAt == null ? c : c.charCodeAt(0);
      })(grid[arr.length - 1 - i][j]) == '*'.charCodeAt(0);
      j++
    )
      putt++;
    sorted[i] = putt;
  }
  return sorted;
}
```
