---
id: 5a23c84252665b21eecc7ec4
title: JortSort
challengeType: 1
forumTopicId: 302293
dashedName: jortsort
---

# --description--

jortSort ni zana ya kupanga inayomfanya mtumizi afanye kazi na inahakikisha ufanisi kwa sababu hutalazimika kupanga tena kamwe. Ilianzishwa awali na Jenn "Moneydollars" Schiffer katika mkutano maarufu wa JSConf2014.

jortSort inapaswa kuwa kitendakazi kinachopokea safu moja ya vitu vinavyoweza kulinganishwa kama hoja yake. Kisha inapanga safu hiyo kwa mpangilio wa kuongezeka na kulinganisha safu iliyopangwa na safu ya awali iliyotolewa. Ikiwa safu hizo zinafanana (yaani safu ya awali ilikuwa tayari imepangwa), kitendakazi hurejesha kweli. Ikiwa safu hizo hazifanani (yaani safu ya awali haikuwa imepangwa), kitendakazi hurejesha si kweli.

# --hints--

`jortsort` inapaswa kuwa kitendakazi.

```js
assert(typeof jortsort == 'function');
```

`jortsort([1,2,3,4,5])` inapaswa kurejesha thamani ya kweli au siyo kweli.

```js
assert(typeof jortsort([1, 2, 3, 4, 5]) == 'boolean');
```

`jortsort([1,2,3,4,5])` inapaswa kurejesha `true`.

```js
assert.equal(jortsort([1, 2, 3, 4, 5]), true);
```

`jortsort([1,2,13,4,5])` inapaswa kurejesha `false`.

```js
assert.equal(jortsort([1, 2, 13, 4, 5]), false);
```

`jortsort([12,4,51,2,4])` inapaswa kurejesha `false`.

```js
assert.equal(jortsort([12, 4, 51, 2, 4]), false);
```

`jortsort([1,2])` inapaswa kurejesha `true`.

```js
assert.equal(jortsort([1, 2]), true);
```

`jortsort([5,4,3,2,1])` inapaswa kurejesha `false`.

```js
assert.equal(jortsort([5, 4, 3, 2, 1]), false);
```

`jortsort([1,1,1,1,1])` inapaswa kurejesha `true`.

```js
assert.equal(jortsort([1, 1, 1, 1, 1]), true);
```

# --seed--

## --seed-contents--

```js
function jortsort(array) {

}
```

# --solutions--

```js
function jortsort(array) {
  // sort the array
  var originalArray = array.slice(0);
  array.sort( function(a,b){return a - b} );

  // compare to see if it was originally sorted
  for (var i = 0; i < originalArray.length; ++i) {
    if (originalArray[i] !== array[i]) return false;
  }

  return true;
};
```
