---
id: 5a7dad05be01840e1778a0d1
title: Fractran
challengeType: 1
forumTopicId: 302270
dashedName: fractran
---

# --description--

FRACTRAN ni lugha ya programu ya kipekee yenye uwezo wa Turing iliyobuniwa na mwanahisabati John Horton Conway.

Programu ya FRACTRAN ni orodha iliyopangwa ya sehemu chanya $P = (f_1, f_2, \\ldots, f_m)$, pamoja na ingizo la nambari kamili chanya ya awali $n$.

Programu inaendeshwa kwa kusasisha nambari kamili $n$ kama ifuatavyo:

<ul>
  <li>kwa sehemu ya kwanza, $f_i$, katika orodha ambayo $nf_i$ ni nambari kamili, badilisha $n$ na $nf_i$ ;</li>
  <li>rudia kanuni hii hadi hakuna sehemu katika orodha inayozalisha nambari kamili wakati inazidishwa na $n$, kisha simama.</li>
</ul>

Conway alitoa programu ya namba msingi katika FRACTRAN:

$\\dfrac{17}{91}$, $\\dfrac{78}{85}$, $\\dfrac{19}{51}$, $\\dfrac{23}{38}$, $\\dfrac{29}{33}$, $\\dfrac{77}{29}$, $\\dfrac{95}{23}$, $\\dfrac{77}{19}$, $\\dfrac{1}{17}$, $\\dfrac{11}{13}$, $\\dfrac{13}{11}$, $\\dfrac{15}{14}$, $\\dfrac{15}{2}$, $\\dfrac{55}{1}$

Kuanzia na $n=2$, programu hii ya FRACTRAN itabadilisha $n$ kuwa $15=2\\times (\\frac{15}{2})$, kisha $825=15\\times (\\frac{55}{1})$, ikizalisha mfululizo ufuatao wa nambari kamili:

$2$, $15$, $825$, $725$, $1925$, $2275$, $425$, $390$, $330$, $290$, $770$, $\\ldots$

Baada ya 2, mfululizo huu una nguvu zifuatazo za 2:

$2^2=4$, $2^3=8$, $2^5=32$, $2^7=128$, $2^{11}=2048$, $2^{13}=8192$, $2^{17}=131072$, $2^{19}=524288$, $\\ldots$

ambazo ni nguvu za namba msingi 2.

# --instructions--

Andika kitendakazi kinachopokea programu ya fractran kama kigezo cha mfuatano wa herufi na kurudisha nambari 10 za kwanza za programu kama safu ya data. Ikiwa matokeo hayana nambari 10, rudisha nambari hizo kama zilivyo.

# --hints--

`fractran` inapaswa kuwa kitendakazi.

```js
assert(typeof fractran == 'function');
```

`fractran("3/2, 1/3")` inapaswa kurudisha safu ya data.

```js
assert(Array.isArray(fractran('3/2, 1/3')));
```

`fractran("3/2, 1/3")` inapaswa kurudisha `[ 2, 3, 1 ]`.

```js
assert.deepEqual(fractran('3/2, 1/3'), [2, 3, 1]);
```

`fractran("3/2, 5/3, 1/5")` inapaswa kurudisha `[ 2, 3, 5, 1 ]`.

```js
assert.deepEqual(fractran('3/2, 5/3, 1/5'), [2, 3, 5, 1]);
```

`fractran("3/2, 6/3")` inapaswa kurudisha `[ 2, 3, 6, 9, 18, 27, 54, 81, 162, 243 ]`.

```js
assert.deepEqual(fractran('3/2, 6/3'), [2, 3, 6, 9, 18, 27, 54, 81, 162, 243]);
```

`fractran("2/7, 7/2")` inapaswa kurudisha `[ 2, 7, 2, 7, 2, 7, 2, 7, 2, 7 ]`.

```js
assert.deepEqual(fractran('2/7, 7/2'), [2, 7, 2, 7, 2, 7, 2, 7, 2, 7]);
```

`fractran("17/91, 78/85, 19/51, 23/38, 29/33, 77/29, 95/23, 77/19, 1/17, 11/13, 13/11, 15/14, 15/2, 55/1")` inapaswa kurudisha `[ 2, 15, 825, 725, 1925, 2275, 425, 390, 330, 290 ]`.

```js
assert.deepEqual(
  fractran(
    '17/91, 78/85, 19/51, 23/38, 29/33, 77/29, 95/23, 77/19, 1/17, 11/13, 13/11, 15/14, 15/2, 55/1'
  ),
  [2, 15, 825, 725, 1925, 2275, 425, 390, 330, 290]
);
```

# --seed--

## --seed-contents--

```js
function fractran(progStr) {

}
```

# --solutions--

```js
function fractran(progStr){
  var num = new Array();
  var den = new Array();
  var val ;
  var out="";
  function compile(prog){
    var regex = /\s*(\d*)\s*\/\s*(\d*)\s*(.*)/m;
    while(regex.test(prog)){
      num.push(regex.exec(prog)[1]);
      den.push(regex.exec(prog)[2]);
      prog = regex.exec(prog)[3];
    }
  }

  function step(val){
    var i=0;
    while(i<den.length && val%den[i] != 0) i++;
    return num[i]*val/den[i];
  }

  var seq=[]

  function exec(val){
    var i = 0;
    while(val && i<limit){
      seq.push(val)
      val = step(val);
      i ++;
    }
  }

  // Main
  compile(progStr);
  var limit = 10;
  exec(2);
  return seq;
}
```
