---
id: 5e6dee7749a0b85a3f1fc7d5
title: Mtihani wa Lucas-Lehmer
challengeType: 1
forumTopicId: 385281
dashedName: lucas-lehmer-test
---

# --description--

Mtihani wa Lucas-Lehmer: kwa $p$ namba msingi isiyoeleweka, namba ya Mersenne $2^p-1$ ni msingi ikiwa na tu ikiwa $2^p-1$ hugawanya $S(p-1)$ ambapo $S(n+1)=(S(n))^2-2$, na $S(1)=4$.

# --instructions--

Andika kitendakazi kinachorejesha kama namba ya Mersenne iliyotolewa ni msingi au siyo msingi.

# --hints--

`lucasLehmer` inapaswa kuwa kitendakazi.

```js
assert(typeof lucasLehmer == 'function');
```

`lucasLehmer(11)` inapaswa kurejesha thamani ya kweli au siyo kweli.

```js
assert(typeof lucasLehmer(11) == 'boolean');
```

`lucasLehmer(11)` inapaswa kurejesha `false`.

```js
assert.equal(lucasLehmer(11), false);
```

`lucasLehmer(15)` inapaswa kurejesha `false`.

```js
assert.equal(lucasLehmer(15), false);
```

`lucasLehmer(13)` inapaswa kurejesha `true`.

```js
assert.equal(lucasLehmer(13), true);
```

`lucasLehmer(17)` inapaswa kurejesha `true`.

```js
assert.equal(lucasLehmer(17), true);
```

`lucasLehmer(19)` inapaswa kurejesha `true`.

```js
assert.equal(lucasLehmer(19), true);
```

`lucasLehmer(21)` inapaswa kurejesha `false`.

```js
assert.equal(lucasLehmer(21), false);
```

# --seed--

## --seed-contents--

```js
function lucasLehmer(p) {

}
```

# --solutions--

```js
function lucasLehmer(p) {
    function isPrime(p) {
        if (p == 2)
            return true;
        else if (p <= 1 || p % 2 == 0)
            return false;
        else {
            var to = Math.sqrt(p);
            for (var i = 3; i <= to; i += 2)
                if (p % i == 0)
                    return false;
            return true;
        }
    }

    function isMersennePrime(p) {
        if (p == 2)
            return true;
        else {
            var m_p = Math.pow(2, p) - 1
            var s = 4;
            for (var i = 3; i <= p; i++)
                s = (s * s - 2) % m_p
            return s == 0;
        }
    }

    return isPrime(p) && isMersennePrime(p)
}
```
