---
id: 5e6dd14797f5ce267c2f19d0
title: Mfululizo wa tazama-na-sema
challengeType: 1
forumTopicId: 385277
dashedName: look-and-say-sequence
---

# --description--

Mfululizo wa tazama na sema ni mfululizo wa nambari uliofafanuliwa kwa njia ya kurudia.

Ufafanuzi wa mfululizo

<ul><li>Chukua nambari ya desimali</li>
<li><span>Tazama</span> nambari hiyo, ukichanganua kwa macho vikundi vinavyofuata vya tarakimu sawa mfululizoni.</li>
<li><span>Sema</span> nambari hiyo, kutoka kushoto kwenda kulia, kundi kwa kundi; jinsi tarakimu hiyo ilivyo nyingi - ikifuatiwa na tarakimu iliyochanganywa.</li></ul><span> Hii inakuwa nambari inayofuata katika mfululizo.</span>

Mfano:

<ul><li>Kuanzia na nambari 1, una <span>moja</span> 1 ambayo hutengeneza 11</li>
<li>Kuanzia na 11, una <span>mbili</span> 1. Yaani: 21</li>
<li>Kuanzia na 21, una <span>moja</span> 2, kisha <span>moja</span> 1. Yaani: (12)(11) ambayo inakuwa 1211</li>
<li>Kuanzia na 1211, una <span>moja</span> 1, <span>moja</span> 2, kisha <span>mbili</span> 1. Yaani: (11)(12)(21) ambayo inakuwa 111221</li></ul>

# --instructions--

Andika kitendakazi kinachokubali mfuatano wa herufi kama kigezo, ukachambue, na kurudisha mfuatano wa herufi unaotokana.

# --hints--

`lookAndSay` inapaswa kuwa kitendakazi.

```js
assert(typeof lookAndSay == 'function');
```

`lookAndSay("1")` inapaswa kurudisha mfuatano wa herufi.

```js
assert(typeof lookAndSay('1') == 'string');
```

`lookAndSay("1")` inapaswa kurudisha `"11"`.

```js
assert.equal(lookAndSay('1'), '11');
```

`lookAndSay("11")` inapaswa kurudisha `"21"`.

```js
assert.equal(lookAndSay('11'), '21');
```

`lookAndSay("21")` inapaswa kurudisha `"1211"`.

```js
assert.equal(lookAndSay('21'), '1211');
```

`lookAndSay("1211")` inapaswa kurudisha `"111221"`.

```js
assert.equal(lookAndSay('1211'), '111221');
```

`lookAndSay("3542")` inapaswa kurudisha `"13151412"`.

```js
assert.equal(lookAndSay('3542'), '13151412');
```

# --seed--

## --seed-contents--

```js
function lookAndSay(str) {

}
```

# --solutions--

```js
function lookAndSay(str) {
    return str.replace(/(.)\1*/g, function(seq, p1) {
      return seq.length.toString() + p1;
    });
}
```
