---
id: 5e94a54cc7b022105bf0fd2c
title: Mzunguko wa maneno
challengeType: 1
forumTopicId: 393913
dashedName: word-frequency
---

# --description--

Ukizingatia mfuatano wa herufi na nambari kamili n, rudisha maneno n yanayotokea mara nyingi zaidi katika faili (pamoja na idadi ya mara yanayotokea) kwa mzunguko unaopungua.

# --instructions--

Andika kitendakazi cha kuhesabu mara za kila neno na kurudisha maneno n yanayotokea mara nyingi zaidi pamoja na idadi ya mara wanazojirudia kwa mzunguko unaopungua.

Kitendakazi kinapaswa kurudisha safu ya data ya vipengele viwili kwa kila kipengele katika muundo ufuatao: `[word, freq]`. `word` inapaswa kuwa toleo la herufi ndogo la neno na `freq` ni nambari inayowakilisha hesabu.

Kitendakazi kinapaswa kurudisha safu ya data tupu, ikiwa hakuna mfuatano wa herufi uliotolewa.

Kitendakazi kinapaswa kutotambua tofauti za herufi kubwa na ndogo, kwa mfano, mfuatano wa herufi "Hello" na "hello" unapaswa kutendewa sawa.

Unaweza kutambua maneno yenye herufi maalum kama vile alama za chini, vibonye, alama za nukuu, koma, n.k., kama maneno tofauti.

Kwa mfano, ukizingatia mfuatano wa herufi "Hello hello goodbye", kitendakazi chako kinapaswa kurudisha `[['hello', 2], ['goodbye', 1]]`.

# --hints--

`wordFrequency` inapaswa kuwa kitendakazi.

```js
assert.isFunction(wordFrequency);
```

`wordFrequency` inapaswa kurudisha safu ya data.

```js
assert.isArray(wordFrequency('test'));
```

`wordFrequency("Hello hello world", 2)` inapaswa kurudisha `[['hello', 2], ['world', 1]]`

```js
const example_1 = 'Hello hello world';
const example_1_solution = [['hello', 2], ['world', 1]];

assert.deepEqual(wordFrequency(example_1, 2), example_1_solution);
```

`wordFrequency("The quick brown fox jumped over the lazy dog", 1)` inapaswa kurudisha `[['the', 2]]`

```js
const example_2 = 'The quick brown fox jumped over the lazy dog';
const example_2_solution = [['the', 2]];

assert.deepEqual(wordFrequency(example_2, 1), example_2_solution);
```

`wordFrequency("Opensource opensource open-source open source", 1)` inapaswa kurudisha `[['opensource', 2]]`

```js
const example_3 = 'Opensource opensource open-source open source';
const example_3_solution = [['opensource', 2]];

assert.deepEqual(wordFrequency(example_3, 1), example_3_solution);
```

`wordFrequency("Apple App apply aPP aPPlE", 3)` inapaswa kurudisha `[['app', 2], ['apple', 2], ['apply', 1]]` au `[['apple', 2], ['app', 2], ['apply', 1]]`

```js
const example_4 = 'Apple App apply aPP aPPlE';
const example_4_solution = [["app", 2],["apple", 2],["apply", 1]];

const arr = wordFrequency(example_4, 3);
assert.includeDeepMembers(arr, example_4_solution);
```

`wordFrequency("c d a d c a b d d c", 4)` inapaswa kurudisha `[['d', 4], ['c', 3], ['a', 2], ['b', 1]]`

```js
const example_5 = 'c d a d c a b d d c';
const example_5_solution = [['d', 4], ['c', 3], ['a', 2], ['b', 1]];

assert.deepEqual(wordFrequency(example_5, 4), example_5_solution);
```

`wordFrequency("", 5)` inapaswa kurudisha `[]`

```js
const example_6 = '';
const example_6_solution = [];

assert.deepEqual(wordFrequency(example_6, 5), example_6_solution);
```

# --seed--

## --seed-contents--

```js
function wordFrequency(txt, n) {

}
```

# --solutions--

```js
function wordFrequency(txt, n) {
  var words = txt.split(/\s+/);
  var wordCount = {};
  words.forEach(word => {
    if (word == '') {
      return;
    }
    const lowerWord = word.toLowerCase();
    wordCount[lowerWord] =
      lowerWord in wordCount ? wordCount[lowerWord] + 1 : 1;
  });

  var wordsArray = [];
  for (let [word, count] of Object.entries(wordCount)) {
    wordsArray.push([word, count]);
  }

  wordsArray.sort((a, b) => {
    if (a[1] !== b[1]) {
      return b[1] - a[1];
    } else if (a[0] !== b[0]) {
      return a[0] < b[0] ? -1 : 1;
    }
    return 0;
  });
  return wordsArray.slice(0, n);
}
```
