---
id: 5e9ddb06ec35240f39657419
title: FizzBuzz
challengeType: 1
forumTopicId: 385370
dashedName: fizzbuzz
---

# --description--

Andika programu inayozalisha safu ya data ya nambari kamili kutoka 1 hadi 100 (ikiwa ni pamoja na). Lakini:

<ul>
    <li>kwa nambari zinazogawika kwa 3, ongeza <code>"Fizz"</code> kwenye safu ya data badala ya nambari</li>
    <li>kwa nambari zinazogawika kwa 5, ongeza <code>"Buzz"</code> kwenye safu ya data badala ya nambari</li>
    <li>kwa nambari zinazogawika kwa 3 na 5, ongeza <code>"FizzBuzz"</code> kwenye safu ya data badala ya nambari</li>
</ul>

# --instructions--

Programu yako inapaswa kurudisha safu ya data yenye matokeo kulingana na sheria zilizo hapo juu.

# --hints--

`fizzBuzz` inapaswa kuwa kitendakazi.

```js
assert(typeof fizzBuzz == 'function');
```

`fizzBuzz()` inapaswa kurudisha safu ya data.

```js
assert(Array.isArray(fizzBuzz()) == true);
```

Nambari zinazogawika kwa 3 pekee zinapaswa kurudisha `"Fizz"`.

```js
assert.equal(fizzBuzz()[2], 'Fizz');
```

Nambari zinazogawika kwa 5 pekee zinapaswa kurudisha `"Buzz"`.

```js
assert.equal(fizzBuzz()[99], 'Buzz');
```

Nambari zinazogawika kwa 3 na 5 zinapaswa kurudisha `"FizzBuzz"`.

```js
assert.equal(fizzBuzz()[89], 'FizzBuzz');
```

Nambari ambazo hazigawiki kwa 3 wala 5 zinapaswa kurudisha nambari yenyewe.

```js
assert.equal(fizzBuzz()[12], 13);
```

# --seed--

## --seed-contents--

```js
function fizzBuzz() {

}
```

# --solutions--

```js
function fizzBuzz() {
    let res=[];
    for (let i =1; i < 101; i++) {
        if (i % 3 === 0  && i % 5 === 0) {
            res.push("FizzBuzz");
        }
        else if (i % 3 === 0) {
            res.push("Fizz");
        }
        else if (i % 5 === 0) {
            res.push("Buzz");
        } 
        else {
            res.push(i);
        }
    }
    return res;
}
```
