---
id: 5eb3e4a21f462f409d656c73
title: Mfululizo wa kujirejea
challengeType: 1
forumTopicId: 385317
dashedName: self-referential-sequence
---

# --description--

Kuna njia kadhaa za kuzalisha mfululizo wa kujirejea. Njia moja inayotumika sana (mfululizo wa <a href="https://rosettacode.org/wiki/Look-and-say_sequence" target="_blank" rel="noopener noreferrer nofollow">Look-and-say</a>) ni kuanza na nambari kamili chanya, kisha kuzalisha kipengele kinachofuata kwa kuunganisha makundi yaliyohesabiwa ya tarakimu zinazofanana zilizo karibu:

<pre>0, 10, 1110, 3110, 132110, 1113122110, 311311222110 ...</pre>

Vipengele vinavyotengenezwa huongezeka kwa urefu kwa kiwango cha kijiometri na havijawahi kuungana.

Njia nyingine ya kuzalisha mfululizo wa kujirejea ni kufupisha kipengele kilichotangulia.

Hesabu ni tarakimu ngapi za kila aina zipo, kisha unganisha jumla na tarakimu kwa kila tarakimu iliyopangwa kwa mpangilio wa nambari. Kumbuka kwamba vipengele vitano vya kwanza ni sawa na vile vya mfululizo uliotangulia.

<pre>0, 10, 1110, 3110, 132110, 13123110, 23124110 ...</pre>

Panga tarakimu kutoka kubwa kwenda ndogo. Usijumlishe hesabu za tarakimu ambazo hazikuonekana katika kipengele kilichotangulia.

Kulingana na thamani ya mbegu, mfululizo unaozalishwa kwa njia hii daima huungana kwa thamani thabiti au kwa mfululizo mfupi wa mizunguko. (Kwa madhumuni yetu, kuungana kunamaanisha kipengele kinafanana na kipengele kilichoshuhudiwa awali.) Mfululizo ulioonyeshwa, ukiwa na thamani ya mbegu 0, huungana kwa thamani thabiti ya 1433223110 baada ya mizunguko 11. Thamani ya mbegu inayoungana haraka zaidi ni 22. Huenda thabiti baada ya kipengele cha kwanza. (Kipengele kinachofuata ni 22, ambacho kimeshawahi kuonekana.)

# --instructions--

Andika kitendakazi kinachopokea thamani ya mbegu kama kigezo, kinazalisha mfululizo wa kujirejea hadi uungane, na kurudisha kama safu ya data.

# --hints--

`selfReferential` inapaswa kuwa kitendakazi.

```js
assert(typeof selfReferential === 'function');
```

`selfReferential(40)` inapaswa kurudisha safu ya data.

```js
assert(Array.isArray(selfReferential(40)));
```

`selfReferential(40)` inapaswa kurudisha `["40", "1410", "142110", "14123110", "1413124110", "2413125110", "151413224110", "152413225110", "251413324110", "152423224110", "152413423110"]`.

```js
assert.deepEqual(selfReferential(40), [
  '40',
  '1410',
  '142110',
  '14123110',
  '1413124110',
  '2413125110',
  '151413224110',
  '152413225110',
  '251413324110',
  '152423224110',
  '152413423110'
]);
```

`selfReferential(132110)` inapaswa kurudisha `["132110", "13123110", "23124110", "1413223110", "1423224110", "2413323110", "1433223110"]`.

```js
assert.deepEqual(selfReferential(132110), [
  '132110',
  '13123110',
  '23124110',
  '1413223110',
  '1423224110',
  '2413323110',
  '1433223110'
]);
```

`selfReferential(132211)` inapaswa kurudisha `["132211", "132231", "232221", "134211", "14131231", "14231241", "24132231", "14233221"]`.

```js
assert.deepEqual(selfReferential(132211), [
  '132211',
  '132231',
  '232221',
  '134211',
  '14131231',
  '14231241',
  '24132231',
  '14233221'
]);
```

`selfReferential(1413223110)` inapaswa kurudisha `["1413223110", "1423224110", "2413323110", "1433223110"]`.

```js
assert.deepEqual(selfReferential(1413223110), [
  '1413223110',
  '1423224110',
  '2413323110',
  '1433223110'
]);
```

`selfReferential(251413126110)` inapaswa kurudisha `["251413126110", "16151413225110", "16251413226110", "26151413325110", "16251423225110", "16251413424110", "16153413225110"]`.

```js
assert.deepEqual(selfReferential(251413126110), [
  '251413126110',
  '16151413225110',
  '16251413226110',
  '26151413325110',
  '16251423225110',
  '16251413424110',
  '16153413225110'
]);
```

# --seed--

## --seed-contents--

```js
function selfReferential(n) {

}
```

# --solutions--

```js
function selfReferential(n) {
  var descending,
    i,
    incr,
    j,
    max_i,
    max_len,
    max_seq,
    seq,
    sequence,
    indexOf =
      [].indexOf ||
      function(item) {
        for (var i = 0, l = this.length; i < l; i++) {
          if (i in this && this[i] === item) return i;
        }
        return -1;
      };

  sequence = function(n) {
    var c, cnt, cnts, d, digit, i, j, l, len, new_cnts, ref, s, seq;
    cnts = {};
    ref = n.toString();
    for (j = 0, len = ref.length; j < len; j++) {
      c = ref[j];
      d = parseInt(c);
      incr(cnts, d);
    }
    seq = [ref];
    while (true) {
      s = '';
      for (i = l = 9; l >= 0; i = --l) {
        if (cnts[i]) {
          s += '' + cnts[i] + i;
        }
      }
      if (indexOf.call(seq, s) >= 0) {
        break;
      }
      seq.push(s);
      new_cnts = {};
      for (digit in cnts) {
        cnt = cnts[digit];
        incr(new_cnts, cnt);
        incr(new_cnts, digit);
      }
      cnts = new_cnts;
    }
    return seq;
  };

  incr = function(h, k) {
    if (h[k] == null) {
      h[k] = 0;
    }
    return (h[k] += 1);
  };

  descending = function(n) {
    var tens;
    if (n < 10) {
      return true;
    }
    tens = n / 10;
    if (n % 10 > tens % 10) {
      return false;
    }
    return descending(tens);
  };

  return sequence(n);
}
```
