---
id: 5eb3e497b8d6d7f63c5517ea
title: Tafuta katika orodha ya rekodi
challengeType: 1
forumTopicId: 385315
dashedName: search-a-list-of-records
---

# --description--

Rekodi ina sifa za kipengele zinazomwelezea kiumbe. Kila sifa ya kipengele ina jina na thamani. Kwa mfano, mtu anaweza kuwa na sifa ya kipengele `age` yenye thamani ya 25. Operesheni muhimu kwenye orodha ya rekodi ni kupata rekodi yenye thamani fulani ya sifa ya kipengele.

# --instructions--

Andika kitendakazi kinachopokea mfuatano wa herufi kama kigezo. Kitendakazi kinapaswa kurudisha kielezo cha kipengee katika `list` ambacho thamani ya sifa ya kipengele `name` inalingana na mfuatano wa herufi uliotolewa.

# --hints--

`searchCity` inapaswa kuwa kitendakazi.

```js
assert(typeof searchCity === 'function');
```

`searchCity("Dar Es Salaam")` inapaswa kurudisha nambari.

```js
assert(typeof searchCity('Dar Es Salaam') === 'number');
```

`searchCity("Dar Es Salaam")` inapaswa kurudisha `6`.

```js
assert.equal(searchCity('Dar Es Salaam'), 6);
```

`searchCity("Casablanca")` inapaswa kurudisha `9`.

```js
assert.equal(searchCity('Casablanca'), 9);
```

`searchCity("Cairo")` inapaswa kurudisha `1`.

```js
assert.equal(searchCity('Cairo'), 1);
```

`searchCity("Mogadishu")` inapaswa kurudisha `4`.

```js
assert.equal(searchCity('Mogadishu'), 4);
```

`searchCity("Lagos")` inapaswa kurudisha `0`.

```js
assert.equal(searchCity('Lagos'), 0);
```

# --seed--

## --seed-contents--

```js
function searchCity(name) {

}

const list = [
  { name: 'Lagos', population: 21.0 },
  { name: 'Cairo', population: 15.2 },
  { name: 'Kinshasa-Brazzaville', population: 11.3 },
  { name: 'Greater Johannesburg', population: 7.55 },
  { name: 'Mogadishu', population: 5.85 },
  { name: 'Khartoum-Omdurman', population: 4.98 },
  { name: 'Dar Es Salaam', population: 4.7 },
  { name: 'Alexandria', population: 4.58 },
  { name: 'Abidjan', population: 4.4 },
  { name: 'Casablanca', population: 3.98 }
];
```

# --solutions--

```js
function searchCity(name) {
  return list.findIndex(item => item.name === name);
}

const list = [
  { name: 'Lagos', population: 21.0 },
  { name: 'Cairo', population: 15.2 },
  { name: 'Kinshasa-Brazzaville', population: 11.3 },
  { name: 'Greater Johannesburg', population: 7.55 },
  { name: 'Mogadishu', population: 5.85 },
  { name: 'Khartoum-Omdurman', population: 4.98 },
  { name: 'Dar Es Salaam', population: 4.7 },
  { name: 'Alexandria', population: 4.58 },
  { name: 'Abidjan', population: 4.4 },
  { name: 'Casablanca', population: 3.98 }
];
```
