---
id: 64dc916138dbe03bdfe6cec6
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

El objetivo de la Torre de Hanoi es mover todos los discos a la última varilla. Para hacerlo, debes seguir tres reglas simples:

1. Puedes mover solo los discos más arriba
2. Puedes mover solo un disco a la vez
3. No puedes colocar discos más grandes encima de discos más pequeños

Debajo de tu código existente, declara una función vacía llamada `move`. Más adelante, usarás esa función para mover los discos entre las varillas. Por ahora, para evitar errores, usa la palabra clave `pass` dentro del cuerpo de la función.

# --hints--

Debes declarar una función vacía llamada `move`. Recuerda usar la palabra clave `pass` dentro del cuerpo de la función con la indentación correcta.

```js
({ test: () => {
		assert(runPython(`
        import inspect
        inspect.isfunction(move)
    `))
	} 
})
```

# --seed--

## --seed-contents--

```py
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--
```
