---
id: 64dc90949ff85e3b37be40d0
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

La función `range()` devuelve una secuencia inmutable de números. Como puedes ver, el tipo de dato de `rods['A']` es `range`, pero quieres que sea una lista.

Pasa tu llamada a `range()` a la función `list()` para hacer eso.

# --hints--

Debes pasar tu llamada `range()` a la función `list()`.

```js
({ test: () => assert(runPython(`
  rods['A'] == list(range(3, 0, -1))
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
