---
id: 64dc9004f658183afa85f4f1
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Ahora verifica el tipo de dato de tu clave `'A'` pas치ndola a la funci칩n `type()` e imprime el resultado en la terminal.

# --hints--

Debes pasar tu tecla `'A'` a la funci칩n `type()`.

```js
({ test: () => assert.match(code, /type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)/) })
```

Debes imprimir el tipo de `rods['A']` usando la funci칩n `print()`.

```js
({ test: () => assert.match(code, /print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
--fcc-editable-region--
```
