---
id: 64dc922df2919e3c38cead61
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

Al inicio de tu código, declara una variable llamada `NUMBER_OF_DISKS` para almacenar el número de discos y asígnale el valor de `3`. Luego, reemplaza el primer argumento pasado a la función `range()` con tu nueva variable.

# --hints--

Debes declarar una variable llamada `NUMBER_OF_DISKS`.

```js
({test: () => assert(__userGlobals.has('NUMBER_OF_DISKS')) })
```

La variable `NUMBER_OF_DISKS` debe tener el valor `3`.

```js
({test: () => assert.equal(__userGlobals.get('NUMBER_OF_DISKS'), 3) })
```

Debes cambiar el valor del primer argumento de la función `range()` a la variable `NUMBER_OF_DISKS`.

```js
({test: () => assert.match(code, /('|")A\1\s*:\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

def move():
    pass
```
