---
id: 64dc8b1faf9c60304ca8a9b6
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este proyecto, resolverás el rompecabezas matemático conocido como la Torre de Hanoi. El rompecabezas consiste en tres varillas y un número de discos de diferentes diámetros.

El objetivo de este rompecabezas es mover los discos de la primera varilla a la tercera varilla, siguiendo reglas específicas que restringen colocar un disco más grande encima de uno más pequeño.

Comienza creando un diccionario vacío llamado `rods` para representar las varillas. 

# --hints--

Debes tener una variable llamada `rods`.

```js
({ test: () => assert(__userGlobals.has('rods')) })
```

Tu variable `rods` debería ser un diccionario vacío.

```js
({ test: () => assert(runPython(`
    rods == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
