---
id: 64dcdee07a14f0299ec30559
title: Paso 22
challengeType: 20
dashedName: step-22
---

# --description--

Escribiste el código para encontrar el movimiento permitido entre las varillas, pero el movimiento real podría ser en ambas direcciones.  
Después de la llamada a `print()`, declara una variable llamada `forward` y asígnale `False`. Usarás esa variable para verificar en qué dirección necesitas mover el disco entre las varillas.

# --hints--

Debes declarar una variable llamada `forward`.

```js
({ test: () => assert.match(code, /forward\s*=\s*./) })
```

Tu variable `forward` debe estar configurada en `False`.

```js
({ test: () => assert.match(code, /forward\s*=\s*False/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
--fcc-editable-region--
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
--fcc-editable-region--
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
