---
id: 64dc97005bc7943e2978df0a
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Ahora puedes eliminar tu llamada a `print()`. Luego, dentro de la función `move()`, elimina la palabra clave `pass` y muestra el contenido de tu diccionario `rods`.

# --hints--

No deberías tener `print(number_of_moves)` en tu código.

```js
({ test: () => assert.isFalse(/print\s*\(\s*number_of_moves\s*\)/.test(code)) })
```

Debes eliminar la palabra clave `pass` e imprimir el diccionario `rods`.

```js
({ test: () => assert.match(code, /def\s+move\s*\(\s*\)\s*:\s+print\s*\(\s*rods\s*\)(?!\s*pass)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
print(number_of_moves)
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
--fcc-editable-region--
```
