---
id: 64dc8ea01436383a88256d7f
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

El rompecabezas comienza con los discos apilados en la primera varilla, en tamaño decreciente, con el disco más pequeño arriba y el disco más grande abajo. Necesitas llenar tu primera lista con números que representen los distintos tamaños de disco.

En lugar de agregar los elementos manualmente a la primera lista, genera una secuencia de números en cuenta regresiva desde `3` hasta `1` usando la función `range()` y asígnala a `rods['A']`. Aquí, `3` representa el disco más grande en la parte inferior de la pila y `1` representa el disco más pequeño en la parte superior de la pila.

La sintaxis es `range(x, y, h)`, donde `x` es el entero inicial (inclusive), `y` es el último entero (no inclusive) y `h` es la diferencia entre un número y el siguiente en la secuencia.

# --hints--

Debes usar la función `range()` para asignar una secuencia de números a `rods['A']`. La sintaxis para llamar a la función `range()` es `range(x, y, h)`.

```js
({ test: () => assert(runPython(`
  type(rods['A']) is range
  `))
})
```

Debes usar la función `range()` para asignar la secuencia de números desde `3` hasta `1` a `rods['A']`.

```js
({ test: () => assert(runPython(`
 rods['A'] == range(3, 0, -1)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': [],
    'B': [],
    'C': []
}
--fcc-editable-region--
```
