---
id: 64dc90e837670a3ba26a1e61
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Ahora que el tipo es `list` como se requiere, puedes eliminar la llamada a `print()`.

# --hints--

Debes eliminar `print(type(rods['A']))` de tu código.

```js
({ test: () => assert.match(code, /\}(?!\s+print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")A\1\s*\]\s*\)\s*\))/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
