---
id: 64dcd3d61c448e2676501f43
title: Paso 20
challengeType: 20
dashedName: step-20
---

# --description--

Cuando el resto del número de movimiento dividido por 3 es igual a 2, se permite el movimiento entre `'A'` y `'B'` (la fuente y las varillas auxiliares).

Agrega una estructura `elif` para eso. Luego, imprime la cadena apropiada si se cumple la condición.

# --hints--

Debes tener una estructura `elif` para ejecutar cuando `remainder == 2`. 

```js
({test: () => assert.match(code, /elif\s+remainder\s*==\s*2:/)})
```

Debes imprimir la cadena `f'Move {i + 1} allowed between {source} and {auxiliary}'`. 

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*source\s*\}\sand\s\{\s*auxiliary\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
