---
id: 64dc92a9718fb53ca3d1bad7
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

El rompecabezas de la Torre de Hanoi se puede resolver en 2<sup>n</sup> - 1 movimientos, donde n es el número de discos. Declara una variable llamada `number_of_moves` y asigna el número total de movimientos a esta variable.

El operador de potencia en Python es `**`.

# --hints--

Debes declarar una variable llamada `number_of_moves`.

```js
({ test: () => assert(__userGlobals.has('number_of_moves')) })
```

El valor de `number_of_moves` debe ser la expresión para calcular el número de movimientos.

```js
({ test: () => {
    assert(runPython(`  
    number_of_moves == 2**NUMBER_OF_DISKS -1      
    `))
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
```
