---
id: 64de1469fccce61940f498c4
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

El diccionario `rods` representará las tres varillas con sus discos. Asígnale las cadenas `'A'`, `'B'` y `'C'` como claves y establece cada una de ellas como una lista vacía.

# --hints--

Tu diccionario `rods` debe tener una clave `'A'`.

```js
({ test: () => assert(runPython(`'A' in rods`)) })
```

`rods['A']` debe ser una lista vacía.

```js
({ test: () => assert(runPython(`
    rods['A'] == []
  `))
})
```

Tu diccionario `rods` debe tener una clave `'B'`.

```js
({ test: () => assert(runPython(`'B' in rods`)) })
```

`rods['B']` debe ser una lista vacía.

```js
({ test: () => assert(runPython(`
    rods['B'] == []
  `))
})
```

Tu diccionario `rods` debe tener una clave `'C'`.

```js
({ test: () => assert(runPython(`'C' in rods`)) })
```

`rods['C']` debe ser una lista vacía.

```js
({ test: () => assert(runPython(`
    rods['C'] == []
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {}
--fcc-editable-region--
```
