---
id: 64df47b32b92301a815d5ef8
title: Paso 49
challengeType: 20
dashedName: step-49
---

# --description--

Al principio, la llamada recursiva que acabas de agregar se encarga del subproblema de mover `n - 1` discos al segundo poste.

Por esa raz칩n, el argumento `target` corresponde a tu segunda varilla, mientras que el argumento `auxiliary` es la tercera varilla. Ten en cuenta que esas seguir치n intercambi치ndose a medida que avanza la recursi칩n.

Corrige el orden de los argumentos intercambiando `target` y `auxiliary` en tu llamada recursiva.

# --hints--

Debes modificar el orden de los argumentos en tu llamada `move(n - 1, source, auxiliary, target)`.

```js
({ test:  () => assert.match(code, /move\(\s*n\s*-\s*1\s*,\s*source\s*,\s*target\s*,\s*auxiliary\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, auxiliary, target)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
