---
id: 64df496c6a8ddf1b38db1ed6
title: Pasos 50
challengeType: 20
dashedName: step-50
---

# --description--

En un paso anterior, escribiste el código para mover el disco más grande del subproblema a la varilla objetivo.

Ahora, todo lo que necesitas hacer es agregar otra llamada recursiva para mover los `n - 1` discos que ya has desplazado. Copia la primera llamada recursiva y pégala al final del bloque `if`.

Ten en cuenta que los argumentos de la función no están en el orden correcto. Intenta descubrir el orden correcto.

# --hints--

Debes llamar a `move(n - 1,  auxiliary, source, target)` al final de tu función `move`.

```js
({ test: () => assert.match(code, /^\s{8}move\(\s*n\s*-\s*1\s*,\s*auxiliary\s*,\s*source\s*,\s*target\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
