---
id: 64df3f1011888113fbd3d81b
title: Paso 44
challengeType: 20
dashedName: step-44
---

# --description--

Para resolver el rompecabezas con recursión, lo primero que debes hacer es descomponer el problema original en subproblemas más pequeños.

La configuración final con `n` discos apilados en la tercera varilla en orden decreciente se puede obtener moviendo:

- `n - 1` discos desde la varilla origen hasta la varilla auxiliar
- el disco más grande desde el origen hasta el destino
- y luego los `n - 1` discos de la varilla auxiliar a la varilla objetivo.

Entonces, lo primero que la función `move` debe hacer es llamarse a sí misma con `n - 1` como primer argumento. Pero si intentas hacerlo sin definir un caso base, obtendrás un `RecursionError`. Esto sucede porque la función sigue llamándose a sí misma indefinidamente.

Antes de tu comentario y tu llamada a `print()`, agrega la llamada recursiva a la función con `n - 1` como el primer argumento y asegúrate de que el cuerpo de la función se ejecute solo cuando `n` sea mayor que cero. Por ahora, deja los otros argumentos en el mismo orden.

# --hints--

El cuerpo de la función `move` debe comenzar con una sentencia `if` que se active cuando `n` sea mayor que cero.

```js
({ test: () => assert.match(code, /def\s+move\(\s*n\s*,\s*source\s*,\s*auxiliary\s*,\s*target\s*\)\s*:\s+if\s+n\s*>\s*0\s*:/) })
```

Debes llamar a `move(n - 1, source, auxiliary, target)` dentro de la sentencia `if`. Recuerda indentar tu llamada a `print()`.

```js
const transformedCode = "\n" + code.replace(/\r/g, '');
const move = __helpers.python.getDef(transformedCode, "move");
const { function_parameters, function_body } = move;

assert.match(function_parameters, /n\s*,\s*source\s*,\s*auxiliary\s*,\s*target/);
const functionBodyIndent = function_body.match(/ +/)[0];
const re = new RegExp(`${functionBodyIndent}if\\s+n\\s*>\\s*0\\s*:\\s+^${functionBodyIndent}( +)move\\(\\s*n\\s*-\\s*1\\s*,\\s*source\\s*,\\s*auxiliary\\s*,\\s*target\\s*\\).+?^${functionBodyIndent}\\1print\\s*\\(\\s*rods\\s*,\\s*("|')\\\\n\\2\\s*\\)`, "ms");
assert.match(function_body, re);
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods, '\n')
    
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
