---
id: 650c6082e5586f9e3acfcd3b
title: Paso 52
challengeType: 20
dashedName: step-52
---

# --description--

Ahora, elimina el diccionario `rods` y convierte sus claves en las variables `A`, `B` y `C`, respectivamente, manteniendo sus valores. Refactoriza tu código para reflejar estos cambios. Si ves la salida en el terminal, hiciste correctamente este paso. 

# --hints--

Debes convertir `rods['A']` en una variable llamada `A` manteniendo su valor.

```js
assert.match(code, /A\s*=\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/);
```

Debes convertir `rods['B']` en una variable llamada `B` manteniendo su valor.

```js
assert.match(code, /B\s*=\s*\[\s*\]/);
```

Debes convertir `rods['C']` en una variable llamada `C` manteniendo su valor.

```js
assert.match(code, /C\s*=\s*\[\s*\]/);
```

Debes modificar tu llamada a `print()` para imprimir `A`, `B`, `C`, en lugar del objeto `rods`. Mantén el carácter de nueva línea en la llamada a `print()`.

```js
assert.match(code, /print\s*\(\s*A\s*,\s*B\s*,\s*C\s*,\s*('|")\\n\1\s*\)/);
```

Debes tener `target.append(source.pop())` en tu código.

```js
assert.match(code, /target\.append\s*\(\s*source\.pop\s*\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
--fcc-editable-region--
```

