---
id: 587d8259367417b2b2512c85
title: 선택 정렬 알고리즘 구현하기
challengeType: 26
dashedName: implement-the-selection-sort-algorithm
---

# --description--

이 실습에서는 선택 정렬 알고리즘을 구현할 것입니다. 선택 정렬은 리스트에서 최솟값을 선택해 리스트의 첫 번째 값과 교환하는 방식으로 작동합니다. 그런 다음 두 번째 위치에서 시작해 남은 리스트 중 가장 작은 값을 선택하고 두 번째 요소와 교환합니다. 리스트 끝에 도달할 때까지 이 과정을 반복하며 요소를 교환합니다. 이렇게 하면 리스트가 정렬됩니다. 선택 정렬은 모든 경우에 시간 복잡도가 이차적입니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `selectionSort` 함수가 있어야 합니다.
1. `selectionSort` 함수는 배열을 인수로 받아 같은 요소를 포함하지만 최솟값에서 최댓값 순으로 정렬된 배열을 반환해야 합니다.

# --hints--

`selectionSort`는 함수여야 합니다.

```js
assert.isFunction(selectionSort);
```

`selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])`는 순서만 바뀌고 나머지는 변경되지 않은 배열을 반환해야 합니다.

```js
assert.sameMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`selectionSort`는 정렬된 배열(최솟값에서 최댓값 순)을 반환해야 합니다.

```js
assert.sameOrderedMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
)
```

`selectionSort`는 내장 `.sort()` 메서드를 사용하지 않아야 합니다.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
