---
id: 66f53dc2c5bd6a11d6c3282f
title: Construir una aplicación generadora de contraseñas
challengeType: 26
dashedName: lab-password-generator
---

# --description--

En este laboratorio, practicarás usando funciones construyendo un generador de contraseñas aleatorias.

**Objetivo:** Cumple con las historias de usuario a continuación y pasa todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes crear una función llamada `generatePassword` que tome un parámetro, indicando la longitud de la contraseña generada. Puedes nombrar el parámetro como prefieras.
2. Tu función debe devolver una cadena que representa una contraseña generada aleatoriamente. Debes usar la siguiente cadena y diferentes métodos `Math` para ayudarte a devolver una nueva cadena con caracteres aleatorios: `ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()`.
3. Debes definir una variable llamada `password` y asignarle el resultado de llamar a la función `generatePassword` con un argumento numérico que representa la longitud deseada de la contraseña.
4. Debes tener un `console.log` que registre una sola cadena hecha al concatenar el mensaje `Contraseña generada:` y la variable `password` separada por un espacio.

# --hints--

Debes tener una función `generatePassword` con un parámetro. Puedes nombrar el parámetro como prefieras.

```js
assert.isFunction(generatePassword);
assert.lengthOf(generatePassword, 1);
```

Tu función `generatePassword` debe devolver una cadena.

```js
const result = generatePassword(5);
assert.isString(result);
```

Tu función `generatePassword` debe devolver una nueva cadena que tenga la longitud correcta.

```js
const length = 8;
const password = generatePassword(length);
assert.lengthOf(password, length);
```

Tu función debe devolver una contraseña generada aleatoriamente con caracteres válidos.

```js
const validChars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
const password = generatePassword(10);
for (let char of password) {
    assert.include(validChars, char);
}
```

Tu función debe devolver una nueva cadena aleatoria cada vez que se llame.

```js
const password1 = generatePassword(10);
const password2 = generatePassword(10);
assert.notStrictEqual(password1, password2);
```

Debes tener una variable `password`.

```js
assert.isDefined(password);
```

Tu variable `password` debe ser una cadena.

```js
assert.isString(password);
```

Debes llamar a la función `generatePassword` con un argumento numérico y almacenar la contraseña devuelta en la variable `password`.

```js
assert.isDefined(password);
assert.isString(password);
assert.match(password, /^[A-Za-z0-9!@#$%^&*()]+$/);
const length2 = 10;
const password3 = generatePassword(length2);
assert.lengthOf(password3, length2);
assert.equal(password3.length, generatePassword(length2).length);
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+password\s*=\s*generatePassword\(\d+\)\;?/);
```

Debes registrar una sola cadena que combine `Contraseña generada:` y la `password` separadas por un solo espacio usando `+` o un literal de plantilla.

```js
const cleanCode = __helpers.removeJSComments(code);
const condition1 = /console\.log\(\s*["']Generated\s+password:\s*["']\s*\+\s*password\s*\);?/gm.test(cleanCode);
const condition2 = /console\.log\(\s*`Generated\s+password:\s*\$\{password\}`\s*\);?/gm.test(cleanCode);
assert.isTrue(condition1 || condition2);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const generatePassword = (length) => {
    const chars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
    let password = "";
    for (let i = 0; i < length; i++) {
        const randomIndex = Math.floor(Math.random() * chars.length);
        password += chars[randomIndex];
    }
    return password;
};

const password = generatePassword(12);
console.log(`Generated password: ${password}`);
```
