---
id: 66c8ba975ee7230e29f6c4b2
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Agora é hora de testar a função `getLoanMessage`.

Use a tabela abaixo para criar `4` variáveis e seus valores:

| Nome da variável | Valor                        |
| ---------------- | ---------------------------- |
| `duplexLoanMsg`  | `getLoanMessage(85000, 850)` |
| `condoLoanMsg`   | `getLoanMessage(65000, 690)` |
| `carLoanMsg`     | `getLoanMessage(45000, 660)` |
| `noLoanMsg`      | `getLoanMessage(25000, 550)` |

Após isso, registre cada variável no console para ver as mensagens.

Com isso, o projeto de verificação de qualificação de empréstimo está completo!

# --hints--

Você deve ter uma variável chamada `duplexLoanMsg`.

```js
assert.isNotNull(duplexLoanMsg);
```

A variável `duplexLoanMsg` deve ser definida como o resultado de `getLoanMessage(85000, 850)`.

```js
assert.equal(duplexLoanMsg, getLoanMessage(85000, 850));
```

Você deve registrar a variável `duplexLoanMsg` no console.

```js
assert.match(code, /console\.log\(duplexLoanMsg\)/)
```

Você deve ter uma variável chamada `condoLoanMsg`.

```js
assert.isNotNull(condoLoanMsg);
```

A variável `condoLoanMsg` deve ser definida como o resultado de `getLoanMessage(65000, 690)`.

```js
assert.equal(condoLoanMsg, getLoanMessage(65000, 690));
```

Você deve registrar a variável `condoLoanMsg` no console.

```js
assert.match(code, /console\.log\(condoLoanMsg\)/)
```

Você deve ter uma variável chamada `carLoanMsg`.

```js
assert.isNotNull(carLoanMsg);
```

A variável `carLoanMsg` deve ser definida como o resultado de `getLoanMessage(45000, 660)`.

```js
assert.equal(carLoanMsg, getLoanMessage(45000, 660));
```

Você deve registrar a variável `carLoanMsg` no console.

```js
assert.match(code, /console\.log\(carLoanMsg\)/)
```

Você deve ter uma variável chamada `noLoanMsg`.

```js
assert.isNotNull(noLoanMsg);
```

A variável `noLoanMsg` deve ser definida como o resultado de `getLoanMessage(25000, 550)`.

```js
assert.equal(noLoanMsg, getLoanMessage(25000, 550));
```

Você deve registrar a variável `noLoanMsg` no console.

```js
assert.match(code, /console\.log\(noLoanMsg\)/)
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

const duplexLoanMsg = getLoanMessage(85000, 850)
console.log(duplexLoanMsg)

const condoLoanMsg = getLoanMessage(65000, 690)
console.log(condoLoanMsg)

const carLoanMsg = getLoanMessage(45000, 660)
console.log(carLoanMsg)

const noLoanMsg = getLoanMessage(25000, 550)
console.log(noLoanMsg)
```
