---
id: 66c8ba975ee7230e29f6c4af
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Se a renda anual do candidato for maior ou igual a `minIncomeForCondo`, E se a pontuação de crédito for maior ou igual a `minCreditScoreForCondo`, então ele se qualifica para um empréstimo de condomínio e carro.

Verifique se isso é verdade na função `getLoanMessage`. Se for, retorne a string `"You qualify for a condo and car loan."`

# --hints--

A função `getLoanMessage` deve retornar uma string se o candidato se qualificar para um empréstimo de condomínio e carro.

```js
assert.isString(getLoanMessage(45000, 680));
```

A função `getLoanMessage` deve retornar a string `"You qualify for a condo and car loan."`.

```js
assert.strictEqual(getLoanMessage(45000, 680), "You qualify for a condo and car loan.");
```

A função `getLoanMessage` deve retornar `undefined` se a renda anual e a pontuação de crédito do candidato não atenderem aos requisitos para um empréstimo de condomínio.

```js
assert.strictEqual(getLoanMessage(45000, 650), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
