---
id: 67ec2f7924ade2450c06bcd8
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

前のレッスンでは、`useEffect` フックを使ったデータ取得などの副作用の処理について学びました。このワークショップでは、ユーザーの入力に基づいてAPIから果物のデータを取得し、その結果を動的に表示するアプリケーションを作成します。

まず、`FruitsSearch` コンポーネントが `div` の `id` を持つ `search-container` 要素を返すようにしてください。

# --hints--

`div` 要素が必要です。

```js
assert.lengthOf(document.querySelectorAll("div"), 2);
```

`div` 要素には `id` という名前の `search-container` が必要です。

```js
const el = document.getElementById("search-container");
assert.exists(el);
assert.strictEqual(el.tagName, "DIV");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Fruits Search</title>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { FruitsSearch } from './index.jsx';
    ReactDOM.createRoot(document.getElementById('root')).render(<FruitsSearch />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#search-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#search-input {
  padding: 10px;
  width: 80%;
  border: 1px solid #ccc;
  border-radius: 5px;
  margin-bottom: 10px;
}

#results {
  text-align: left;
  max-height: 150px;
  overflow-y: auto;
}
.result-item {
  padding: 5px;
  border-bottom: 1px solid #ddd;
}
```

```jsx
const { useState, useEffect } = React;

--fcc-editable-region--
export function FruitsSearch() {
  
}
--fcc-editable-region--
```
