---
id: 68551b677a5f6898af364b85
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次工作坊中，你將透過使用 Tailwind CSS 的 flexbox 工具類別來建置一個響應式的呼叫行動（CTA）組件，進一步了解 Tailwind CSS。

基本的 HTML 範本已經為你提供，所以你需要先建置 HTML 元素。

首先創建一個 `div` 元素。在該 `div` 內，創建另外兩個獨立的 `div` 元素。

# --hints--

你應該創建一個 `div` 元素。

```js
assert.exists(document.querySelector("div"));
```

你應該在 `div` 元素內再有兩個 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll("div > div"), 2)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
