---
id: 68552464ef378e1051a55c28
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

在第二個 `div` 裡，創建兩個錨點元素。第一個應該有文字 `Learn more`，第二個應該有文字 `Start listening`。

將兩個錨點元素的 `href` 屬性都設為 `#`。

# --hints--

你應該在第二個 `div` 元素內創建兩個 `a` 元素。

```js
assert.lengthOf(document.querySelectorAll("div > div:nth-child(2) > a"), 2)
```

你的第一個 `a` 元素應該在第二個 `div` 元素內有文字 `Learn more` 和 `href` 屬性為 `#`。

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[0]?.textContent, "Learn more")
assert.equal(links[0]?.getAttribute("href"), "#")
```

你的第二個 `a` 元素應該在第二個 `div` 元素內有文字 `Start listening` 和 `href` 屬性為 `#`。

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[1]?.textContent, "Start listening")
assert.equal(links[1]?.getAttribute("href"), "#")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
