---
id: 685523c05fb41d0848c4b52f
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在第一個 `div` 元素內，創建一個帶有文字 `Soundflow` 的 `span` 元素、一個帶有文字 `Discover New Music` 的 `h1` 元素，以及一個帶有文字 `Stream your favorite tracks and discover new artists.` 的 `p` 元素。

# --hints--

你應該在第一個 `div` 元素內創建一個包含文字 `Soundflow` 的 `span` 元素。

```js
assert.equal(document.querySelector("div > div:first-child > span")?.textContent, "Soundflow")
```

你應該在第一個 `div` 元素內創建一個 `h1` 元素，並在其中放入文字 `Discover New Music`。

```js
assert.equal(document.querySelector("div > div:first-child > h1")?.textContent, "Discover New Music")
```

你應該在第一個 `div` 元素內創建一個 `p` 元素，並在其中放入文字 `Stream your favorite tracks and discover new artists.`。

```js
assert.equal(document.querySelector("div > div:first-child > p")?.textContent, "Stream your favorite tracks and discover new artists.")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </div>
      <div></div>
    </div>
  </body>
</html>
```
