---
id: 68551b677a5f6898af364b85
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次研討會中，你將通過使用 flexbox 實用類來構建響應式號召性用語 (CTA) 組件來了解有關 Tailwind CSS 的更多信息。

已經爲你提供了基本的 HTML 樣板，因此你需要首先構建 HTML 元素。

首先創建一個 `div` 元素。 在 `div` 內部，再創建兩個單獨的 `div` 元素。

# --hints--

你應該創建一個 `div` 元素。

```js
assert.exists(document.querySelector("div"));
```

`div` 元素內應該還有兩個 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll("div > div"), 2)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
