---
id: 6858f90c54647a48ddb8822d
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

此時，內容，也就是外層 `div` 內的兩個 `div` 元素，將會無論螢幕大小如何都堆疊在彼此上方。

那是因為你在包裹的 `div` 元素中有 `flex-col` 類別。若要修正，請使用 `lg:` 前置式將 `flex-row` 類別添加到該 `div`。 

這表示在具有較大螢幕的裝置上，flex 方向將會是 `row`。若要再次看到該效果，請將瀏覽器調整為較大寬度。

# --hints--

你的外層 `div` 元素應該有實用類別 `lg:flex-row`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("lg:flex-row"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
