---
id: 6858ee73ed8106ccb214942a
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

既然所有內容都靠左對齊，你需要透過為 x 軸添加自動邊距來置中整個內容。

將對應的 Tailwind 類別添加到外層的 `div` 元素。要看到該效果，請調整預覽標籤直到整個內容置中。

# --hints--

你的外層 `div` 元素應該有工具類別 `mx-auto`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--    
  </body>
</html>
```
