---
id: 6858eadd3144c89d45e02bc7
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

如果你注意到，外層的 `div` 太貼近頂部，且其中的內容太靠左堆疊，所以你現在應該調整一些間距。

為你的 `div` 添加類別 `mt-8` 以稍微向下推動，並添加 `p-4` 以將文字從左側移開。

# --hints--

你的外層 `div` 元素應該使用工具類別 `mt-8`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mt-8"));
```

你的外層 `div` 元素應該使用工具類別 `p-4`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
