---
id: 6858ef974c8ad5dcd7abb9f1
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

`div` 內的內容可以有更好的間距和排列。至於其餘的間距和排列，你需要使用 flexbox。

為你的 `div` 添加類別 `flex`，使兩個子 `div` 元素並排顯示。

此外，為你的內容添加 `justify-around`，使內容周圍和之間的空間均勻分布，並添加 `items-center`，使整個內容最終在 `div` 中置中。 

# --hints--

你的外層 `div` 元素應該有實用類別 `flex`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex"));
```

你的外層 `div` 元素應該有實用類別 `justify-around`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("justify-around"));
```

你的外層 `div` 元素應該有工具類別 `items-center`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
