---
id: 6858ebf2552046ab3de9f17c
title: 步驟7
challengeType: 0
dashedName: step-7
---

# --description--

此時，你會注意到 `div` 中內容右側的空間太多。 你可以通過增加寬度來最小化這種情況。

將類 `md:w-1/2` 添加到 `div` 以將中型和大型屏幕上的寬度減小到 50%。 也就是說，屏幕分辨率爲 768 像素及以上。

要查看效果，請展開預覽選項卡。

# --hints--

你的外部 `div` 元素應該具有實用類 `md:w-1/2`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("md:w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
