---
id: 6858e6408a0999603b2a4fee
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

為 `h1` 添加類別 `font-bold` 和 `text-4xl`。這將使文字更粗、更大，並且增加行距。

此外，為你的 `h1` 添加類別 `my-4`，以將上下邊距增加 16 像素，或 1rem。

# --hints--

你的 `h1` 元素應該有實用類別 `font-bold`。

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("font-bold"));
```

你的 `h1` 元素應該有工具類別 `text-4xl`。

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("text-4xl"));
```

你的 `h1` 元素應該有工具類別 `my-4`。

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("my-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
         --fcc-editable-region--
        <h1>Discover New Music</h1>
         --fcc-editable-region--
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
