---
id: 6858e7e23bf9e475e14cb931
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

既然 HTML 結構已經建立完成，現在是時候使用 Tailwind 進行適當的樣式設計。

從包裹整個內容的 `div` 元素開始。透過給它一個 `class` 為 `bg-indigo-600` 來設定它的背景顏色。

# --hints--

你的外層 `div` 應該有一個設定為 `bg-indigo-600` 的 `class` 屬性。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("bg-indigo-600"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
   --fcc-editable-region--
  </body>
</html>
```
