---
id: 6858e7e23bf9e475e14cb931
title: 步驟4
challengeType: 0
dashedName: step-4
---

# --description--

現在 HTML 結構已經構建完畢，是時候使用 Tailwind 轉到適當的樣式了。

從包裹整個內容的 `div` 元素開始。 通過賦予 `class` 爲 `bg-indigo-600` 來設置其背景顏色。

# --hints--

你的外部 `div` 應該將 `class` 屬性設置爲 `bg-indigo-600`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("bg-indigo-600"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
   --fcc-editable-region--
  </body>
</html>
```
