---
id: 6858f3a67aa8580ff37016c7
title: 步驟 10
challengeType: 0
dashedName: step-10
---

# --description--

如果你再次將預覽選項卡的寬度調整爲與移動設備相同的寬度，你會注意到內容始終並排放置。 這並不是你在移動設備上想要的結果。

爲了解決這個問題，請將類 `flex-col` 添加到 `div`。 這會將彈性方向改變爲列，因此元素會堆疊在一起。

以下是使用 `flex-col` 實用類的方法：

```html
<div class="flex flex-col">
  <element></element>
  <element></element>
</div>
```

# --hints--

你的起始 `div` 元素應該具有實用類 `flex-col`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex-col"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
