---
id: 6858ebf2552046ab3de9f17c
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

此時，你會注意到 `div` 中內容右側有過多空間。你可以透過添加寬度來將其縮小。

為你的 `div` 添加類別 `md:w-1/2`，以在中等及更大螢幕上將寬度減少到 50%。也就是說，螢幕寬度為 768 像素及以上。

要看到效果，請展開預覽標籤。

# --hints--

你的外層 `div` 元素應該有實用類別 `md:w-1/2`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("md:w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
