---
id: 6858f3a67aa8580ff37016c7
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

如果你將預覽分頁的寬度調整為與行動裝置相同的寬度，你會注意到內容總是並排顯示。這並不是你在行動裝置上想要的效果。

為了修正這個問題，為你的 `div` 添加類別 `flex-col`。這將把彈性方向改為行，因此元素會堆疊在彼此上方。

以下是如何使用 `flex-col` 實用類別：

```html
<div class="flex flex-col">
  <element></element>
  <element></element>
</div>
```

# --hints--

你的開啟 `div` 元素應該有工具類別 `flex-col`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex-col"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
