---
id: 68597114730d7c6b8838d204
title: 第 15 步
challengeType: 0
dashedName: step-15
---

# --description--

`h1` 下方的段落看起來不錯，因此不需要更多的樣式。這表示你可以移動到第二個 `div` 元素內的錨點元素。

為第一個錨點元素添加類別 `bg-white`、`hover:bg-gray-100` 和 `text-indigo-600`。這會給它一個白色背景、靛藍色的文字顏色，並且當你暫留在上面時會有稍微變亮的背景。

為第二個錨點元素添加類別 `bg-pink-500`、`hover:bg-pink-600` 和 `text-white`。

# --hints--

你的第一個錨點元素應該有實用類別 `bg-white`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("bg-white"));
```

你的第一個錨點元素應該有實用類別 `hover:bg-gray-100`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("hover:bg-gray-100"));
```

你的第一個錨點元素應該有實用類別 `text-indigo-600`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("text-indigo-600"));
```

你的第二個錨點元素應該有實用類別 `bg-pink-500`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("bg-pink-500"));
```

你的第二個錨點元素應該有實用類別 `hover:bg-pink-600`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("hover:bg-pink-600"));
```

你的第二個錨點元素應該有實用類別 `text-white`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("text-white"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        <a href="#">
          Learn more
        </a>
        <a href="#">
          Start listening
        </a>
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
