---
id: 587d7fb7367417b2b2512c0a
title: model.create() で複数のレコードを作成する
challengeType: 2
forumTopicId: 301537
dashedName: create-many-records-with-model-create
---

# --description--

時にはモデルのインスタンスを複数作成する必要があります。例えば、データベースに初期データを投入するときなどです。`Model.create()` は `[{name: 'John', ...}, {...}, ...]` のようなオブジェクトの配列を最初の引数に取り、それらをすべてデータベースに保存します。

# --instructions--

`createManyPeople` 関数を修正して、`Model.create()` を使い `arrayOfPeople` を引数にして複数の人を作成してください。

**注意:** 前の演習でインスタンス化したモデルを再利用できます。

# --hints--

一度に複数のデータベースアイテムを作成することが成功するはずです

```js
  const response = await fetch(code + '/_api/create-many-people', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'John', age: 24, favoriteFoods: ['pizza', 'salad'] },
      { name: 'Mary', age: 21, favoriteFoods: ['onions', 'chicken'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an array');
  assert.equal(
    data.length,
    2,
    'the response does not contain the expected number of items'
  );
  assert.equal(data[0].name, 'John', 'The first item is not correct');
  assert.equal(
    data[0].__v,
    0,
    'The first item should be not previously edited'
  );
  assert.equal(data[1].name, 'Mary', 'The second item is not correct');
  assert.equal(
    data[1].__v,
    0,
    'The second item should be not previously edited'
  );
```
