---
id: 587d7fb8367417b2b2512c0f
title: model.findOneAndUpdate() を使ってドキュメントを新しく更新する
challengeType: 2
forumTopicId: 301542
dashedName: perform-new-updates-on-a-document-using-model-findoneandupdate
---

# --description--

Mongoose の最近のバージョンには、ドキュメントの更新を簡単にするメソッドがあります。より高度な機能（例えば、pre/post フックやバリデーション）はこの方法では動作が異なるため、従来の方法も多くの状況でまだ役立ちます。`findByIdAndUpdate()` は id で検索するときに使えます。

# --instructions--

`findAndUpdate` 関数を修正して、`Name` で人を検索し、その人の年齢を `20` に設定してください。検索キーとして関数のパラメータ `personName` を使います。

**注意:** 更新されたドキュメントを返す必要があります。そのためには、オプションのドキュメント `{ new: true }` を `findOneAndUpdate()` の第3引数として渡してください。デフォルトでは、これらのメソッドは変更されていないオブジェクトを返します。

# --hints--

findOneAndUpdate でアイテムを更新することに成功するはずです

```js
  const response = await fetch(code + '/_api/find-one-update', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Dorian Gray',
      age: 35,
      favoriteFoods: ['unknown']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Dorian Gray', 'item.name is not what expected');
  assert.equal(data.age, 20, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['unknown'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(
    data.__v,
    0,
    'findOneAndUpdate does not increment version by design!'
  );
```
