---
id: 587d7fb8367417b2b2512c11
title: model.remove() で複数のドキュメントを削除する
challengeType: 2
forumTopicId: 301538
dashedName: delete-many-documents-with-model-remove
---

# --description--

`Model.remove()` は、指定した条件に一致するすべてのドキュメントを削除するのに便利です。

# --instructions--

`removeManyPeople` 関数を修正して、`nameToRemove` 変数内に名前が含まれるすべての人を `Model.remove()` を使って削除してください。`name` フィールドを設定したクエリドキュメントとコールバックを渡してください。

**注意:** `Model.remove()` は削除されたドキュメントを返さず、操作の結果と影響を受けたアイテム数を含む JSON オブジェクトを返します。テストで使用するため、`done()` コールバックに渡すのを忘れないでください。

# --hints--

一度に複数のアイテムを削除することに成功するはずです

```js
  const response = await fetch(code + '/_api/remove-many-people', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'Mary', age: 16, favoriteFoods: ['lollipop'] },
      { name: 'Mary', age: 21, favoriteFoods: ['steak'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(!!data.ok, 'The mongo stats are not what expected');
  assert.equal(
    data.n,
    2,
    'The number of items affected is not what expected'
  );
  assert.equal(data.count, 0, 'the db items count is not what expected');
```
