---
id: 587d7fb7367417b2b2512c0b
title: モデルの find() を使ってデータベースを検索する
challengeType: 2
forumTopicId: 301543
dashedName: use-model-find-to-search-your-database
---

# --description--

最も基本的な使い方では、`Model.find()` は最初の引数にクエリドキュメント（JSONオブジェクト）を受け取り、次にコールバックを受け取ります。マッチした結果の配列を返します。非常に幅広い検索オプションに対応しています。詳細はドキュメントをご覧ください。

# --instructions--

`findPeopleByName` 関数を修正して、<code>Model.find() -&gt; [Person]</code> を使い、指定された名前を持つすべての人を検索してください。

関数の引数 `personName` を検索キーとして使ってください。

# --hints--

条件に合うすべての項目を検索することに成功するはずです

```js
  const response = await fetch(code + '/_api/find-all-by-name', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'r@nd0mN4m3',
      age: 24,
      favoriteFoods: ['pizza']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an Array');
  assert.equal(
    data[0].name,
    'r@nd0mN4m3',
    'item.name is not what expected'
  );
  assert.equal(data[0].__v, 0, 'The item should be not previously edited');
```
