---
id: 587d7fb7367417b2b2512c0d
title: _id でデータベースを検索するために model.findById() を使う
challengeType: 2
forumTopicId: 301544
dashedName: use-model-findbyid-to-search-your-database-by-id
---

# --description--

ドキュメントを保存すると、MongoDB は自動的に `_id` フィールドを追加し、一意の英数字キーを設定します。`_id` での検索は非常に頻繁に行われる操作なので、Mongoose はそれ専用のメソッドを提供しています。

# --instructions--

`findPersonById` を修正して、`_id` を使い、指定された `Model.findById() -> Person` を持つ唯一の人物を検索してください。検索キーとして関数の引数 `personId` を使います。

# --hints--

Id でアイテムを検索することは成功するはずです

```js
  const response = await fetch(code + '/_api/find-by-id');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'test', 'item.name is not what expected');
  assert.equal(data.age, 0, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['none'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```
