---
id: 587d7fb7367417b2b2512c0c
title: データベースから単一の一致するドキュメントを返すために model.findOne() を使う
challengeType: 2
forumTopicId: 301545
dashedName: use-model-findone-to-return-a-single-matching-document-from-your-database
---

# --description--

`Model.findOne()` は `Model.find()` と似ていますが、複数のアイテムがあっても配列ではなく1つのドキュメントだけを返します。特にユニークとして宣言したプロパティで検索するときに便利です。

# --instructions--

`findOneByFood` 関数を修正して、`Model.findOne() -> Person` を使い、引数 `food` を検索キーとして、その人のお気に入りの食べ物に特定のものがある1人だけを見つけてください。

# --hints--

1つのアイテムを見つけることが成功するはずです

```js
  const response = await fetch(code + '/_api/find-one-by-food', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Gary',
      age: 46,
      favoriteFoods: ['chicken salad']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Gary', 'item.name is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['chicken salad'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```
