---
id: 587d7fb9367417b2b2512c12
title: 検索結果を絞り込むためのチェーン検索クエリヘルパー
challengeType: 2
forumTopicId: 301533
dashedName: chain-search-query-helpers-to-narrow-search-results
---

# --description--

`Model.find()`（または他の検索メソッド）にコールバックを最後の引数として渡さないと、クエリは実行されません。クエリは変数に保存して後で使うこともできます。この種のオブジェクトはチェーン構文を使ってクエリを組み立てることを可能にします。実際のデータベース検索は、最後に`.exec()`メソッドをチェーンしたときに実行されます。この最後のメソッドには必ずコールバックを渡す必要があります。多くのクエリヘルパーがありますが、ここでは最もよく使われるものを使います。

# --instructions--

`queryChain`関数を修正して、変数`foodToSearch`で指定された食べ物が好きな人を見つけてください。`name`でソートし、結果を2件に制限し、年齢を非表示にします。`.find()`、`.sort()`、`.limit()`、`.select()`、そして`.exec()`をチェーンしてください。`done(err, data)`コールバックを`exec()`に渡してください。

# --hints--

チェーンでクエリヘルパーを使うことが成功の鍵です

```js
  const response = await fetch(code + '/_api/query-tools', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'Pablo', age: 26, favoriteFoods: ['burrito', 'hot-dog'] },
      { name: 'Bob', age: 23, favoriteFoods: ['pizza', 'nachos'] },
      { name: 'Ashley', age: 32, favoriteFoods: ['steak', 'burrito'] },
      { name: 'Mario', age: 51, favoriteFoods: ['burrito', 'prosciutto'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an Array');
  assert.equal(
    data.length,
    2,
    'the data array length is not what expected'
  );
  assert.notProperty(
    data[0],
    'age',
    'The returned first item has too many properties'
  );
  assert.equal(
    data[0].name,
    'Ashley',
    'The returned first item name is not what expected'
  );
  assert.notProperty(
    data[1],
    'age',
    'The returned second item has too many properties'
  );
  assert.equal(
    data[1].name,
    'Mario',
    'The returned second item name is not what expected'
  );
```
