---
id: 587d7fb8367417b2b2512c10
title: model.findByIdAndRemove を使ってドキュメントを1件削除する
challengeType: 2
forumTopicId: 301539
dashedName: delete-one-document-using-model-findbyidandremove
---

# --description--

`findByIdAndRemove` と `findOneAndRemove` は前回の更新メソッドと似ています。削除されたドキュメントをデータベースに渡します。いつも通り、関数の引数 `personId` を検索キーとして使ってください。

# --instructions--

`removeById` 関数を修正して、人物の `_id` を使って1人を削除してください。`findByIdAndRemove()` または `findOneAndRemove()` のどちらかのメソッドを使う必要があります。

# --hints--

アイテムの削除が成功するはずです

```js
  const response = await fetch(code + '/_api/remove-one-person', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Jason Bourne',
      age: 36,
      favoriteFoods: ['apples']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Jason Bourne', 'item.name is not what expected');
  assert.equal(data.age, 36, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['apples'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0);
  assert.equal(data.count, 0, 'the db items count is not what expected');
```
