---
id: 587d7fb6367417b2b2512c09
title: モデルのレコードを作成して保存する
challengeType: 2
forumTopicId: 301536
dashedName: create-and-save-a-record-of-a-model
---

# --description--

このチャレンジでは、モデルのレコードを作成して保存する必要があります。

# --instructions--

`createAndSavePerson` 関数内で、以前作成した `Person` モデルのコンストラクタを使ってドキュメントインスタンスを作成してください。コンストラクタには `name`、`age`、`favoriteFoods` のフィールドを持つオブジェクトを渡してください。それらの型は `personSchema` にあるものと一致している必要があります。次に、返されたドキュメントインスタンスに対して `document.save()` メソッドを呼び出してください。Node の慣例に従ったコールバック関数を引数として渡してください。これは一般的なパターンで、以下のすべての CRUD メソッドは最後の引数としてこのようなコールバック関数を取ります。

```js
/* Example */

// ...
person.save(function(err, data) {
  //   ...do your stuff here...
});
```

# --hints--

データベースアイテムの作成と保存は成功するはずです

```js
  const response = await fetch(code + '/_api/create-and-save-person');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isString(data.name, '"item.name" should be a String');
  assert.isNumber(data.age, '28', '"item.age" should be a Number');
  assert.isArray(
    data.favoriteFoods,
    '"item.favoriteFoods" should be an Array'
  );
  assert.equal(data.__v, 0, 'The db item should be not previously edited');
```
