---
id: 656470d517833a39bb8b5608
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

O módulo `random` contém um gerador de números pseudoaleatórios.A maioria das suas funcionalidades depende da função `random()`, que retorna um número de ponto flutuante no intervalo entre `0.0` (inclusivo) e `1.0` (exclusivo).

Chame a função `random()` do módulo `random` e imprima o resultado.

# --hints--

Você deve imprimir o resultado da chamada `random.random()`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.random\s*\(\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--

--fcc-editable-region--
```
