---
id: 6564d27c8e4c197d367c3b1c
title: Passo 39
challengeType: 20
dashedName: step-39
---

# --description--

Adicione uma nova tupla à lista `constraints`.Use `lowercase` como o primeiro item e um padrão regex que corresponde a uma única letra minúscula como o segundo item.

# --hints--

Você deve adicionar uma segunda tupla à lista `constraints`.Use o parâmetro `lowercase` como o primeiro item e `[a-z]` como o segundo item.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*("|')\[0-9\]\1\s*\)\s*,\s*\(\s*lowercase\s*,\s*("|')\[a-z\]\2\s*\)\s*,?\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, '[0-9]')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = '[^a-z]t'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
