---
id: 656469e62377ab34a5bcb8d1
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Declare duas novas variáveis chamadas `digits` e `symbols` e atribua a elas `string.digits` e `string.punctuation`, respectivamente.

# --hints--

Você deve declarar uma variável chamada `digits`.

```js
({ test: () => assert(__userGlobals.has("digits")) })
```

Você deve atribuir `string.digits` à sua variável `digits`.

```js
({ test: () => assert(runPython(`
    import string
    digits == string.digits
  `))
})
```

Você deve declarar uma variável chamada `symbols`.

```js
({ test: () => assert(__userGlobals.has("symbols")) })
```

Você deve atribuir `string.punctuation` à sua variável `symbols`.

```js
({ test: () => assert(runPython(`
    import string
    symbols == string.punctuation
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string

letters = string.ascii_letters
--fcc-editable-region--
```
