---
id: 65646ab6bced52360f44670f
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Agora declare uma variável chamada `all_characters` e atribua a ela o resultado da concatenação das suas variáveis existentes.

# --hints--

Você deve declarar uma variável chamada `all_characters`.

```js
({ test: () => assert(__userGlobals.has("all_characters")) })
```

Você deve concatenar `letters`, `digits` e `symbols` e atribuir o resultado à sua variável `all_characters`.

```js
({ test: () => assert(runPython(`
    import string
    from itertools import permutations
    perms = permutations([string.ascii_letters, string.digits, string.punctuation])
    any("".join(perm) == all_characters for perm in perms)
  `))
})
```

# --seed--

## --seed-contents--

```py
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--

--fcc-editable-region--
```
