---
id: 6564883669b5af4b69f794cc
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Sua variável `all_characters` é uma string formada por todas as letras minúsculas e maiúsculas, todos os 10 dígitos e vários caracteres especiais. 

Logo antes, adicione um comentário dizendo `Combine all characters`.

# --hints--

Você deve adicionar o comentário logo acima da sua variável `all_characters`.

```js
({ test: () => assert.match(code, /^#\s*Combine all characters\s+^all_characters/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation


all_characters = letters + digits + symbols
--fcc-editable-region--
```
