---
id: 6564d97a5ef99783216a6229
title: Passo 44
challengeType: 20
dashedName: step-44
---

# --description--

Python fornece um tipo particular de string chamado string *raw*. Strings raw são prefixadas com `r`. A principal distinção em relação às strings comuns está em como elas tratam o caractere de barra invertida: em strings raw, as barras invertidas são tratadas como caracteres literais e não como caracteres de escape. Ao escrever expressões regulares, usar strings raw é uma boa prática, já que elas geralmente podem conter muitos caracteres `\`.

Transforme sua string `pattern` em uma string raw prefixando-a com um `r`.

# --hints--

Você deve modificar sua variável `pattern` para `r'\.'`.

```js
({ test: () => assert.match(code, /^pattern\s*=\s*r("|')\\\.\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '\.'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
