---
id: 6564dd65c3c2fa873a83d213
title: Passo 47
challengeType: 20
dashedName: step-47
---

# --description--

Em uma classe de caracteres, você pode combinar múltiplos intervalos escrevendo um intervalo após o outro dentro dos colchetes (sem quaisquer caracteres adicionais).Por exemplo: `[a-d3-6]` é a combinação de `[a-d]` e `[3-6]`.

Agora, modifique o último padrão regex para corresponder a qualquer caractere não alfanumérico.Combine os intervalos `a-z`, `A-Z` e `0-9` em uma única classe de caracteres e adicione um `^` como o primeiro caractere para negar o padrão.

# --hints--

Você deve completar seu quarto padrão para corresponder a qualquer caractere não alfanumérico.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*r("|')\\d\1\s*\)\s*,\s*\(\s*lowercase\s*,\s*r("|')\[a-z\]\2\s*\)\s*,\s*\(\s*uppercase\s*,\s*r("|')\[A-Z\]\3\s*\)\s*,\s*\(\s*special_chars\s*,\s*r("|')\[\^a-zA-Z0-9\]\4\s*\)\s*,?\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = r'\.'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
