---
id: 656474314bf37d3c83a85143
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

A função `choice()` do módulo `random` recebe uma sequência e retorna um item aleatório da sequência.

Modifique sua chamada `print()` para usar a função `choice()` e passar `all_characters` como argumento.

# --hints--

Você deve modificar sua chamada existente `print()` para imprimir `random.choice(all_characters)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--
print(random.random())
--fcc-editable-region--
```
