---
id: 6564f32b18480893cf7799fd
title: Passo 52
challengeType: 20
dashedName: step-52
---

# --description--

Agora, combine sua raw string com uma f-string e interpole sua variável `symbols` dentro da classe de caracteres. Lembre-se de que você pode interpolar uma variável dentro de uma f-string usando chaves `{ }`.

# --hints--

O segundo item no seu quarto tuple de restrição deve ser a string `fr'[{symbols}]'`.

```js
({ test: () => assert.match(code, /\(\s*special_chars\s*,\s*(fr|rf)("|')\[\{\s*symbols\s*\}\]\2/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
--fcc-editable-region--            
            (special_chars, r'\W')
--fcc-editable-region--            
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
```
