---
id: 656494269ccce754411a2112
title: Passo 27
challengeType: 20
dashedName: step-27
---

# --description--

O módulo `re` permite que você use *expressões regulares* no seu código.Você aprenderá mais sobre expressões regulares muito em breve.

Por enquanto, adicione uma declaração de `import` no topo do seu código para importar o módulo `re`.

# --hints--

Você deve importar o módulo `re`.

```js
({ test: () => assert.match(code, /^import\s+re/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

import secrets
import string
--fcc-editable-region--

def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
```
