---
id: 656477845006313fbfea0ad1
title: Passo 12
challengeType: 20
dashedName: step-12
---

# --description--

Embora o efeito possa parecer igual ao `random.choice()`, o `secrets` garante a fonte de aleatoriedade mais segura que seu sistema operacional pode fornecer.

Agora, delete suas duas chamadas `print()`.

# --hints--

Você deve deletar suas duas chamadas `print()`.

```js
({ test: () => assert.isFalse( /^print/m.test(code)) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols
--fcc-editable-region--
print(all_characters)
print(secrets.choice(all_characters))
--fcc-editable-region--
```
