---
id: 6564b8c9349bd76dc037967b
title: Passo 29
challengeType: 20
dashedName: step-29
---

# --description--

A função `search()` do módulo `re` analisa a string passada como argumento procurando o primeiro local onde o padrão regex corresponde à string.

Declare uma variável chamada `quote` e atribua a string `'Not all those who wander are lost.'` a essa variável.Então, imprima o resultado de `pattern.search(quote)`.

# --hints--

Você deve ter uma variável `quote`.

```js
({ test: () => assert(__userGlobals.has("quote")) })
```

Você deve atribuir a string fornecida à sua nova variável `quote`.

```js
({ test: () => assert.equal(__userGlobals.get("quote"), "Not all those who wander are lost.") })
```

Você deve imprimir `pattern.search(quote)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*pattern\.search\s*\(\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('i')

--fcc-editable-region--
```
