---
id: 6564cbb6311a40783b3f5de6
title: Passo 36
challengeType: 20
dashedName: step-36
---

# --description--

Classes de caracteres também permitem que você indique um intervalo de caracteres para corresponder.Você precisa especificar os caracteres inicial e final separados por um hífen, `-`.Os caracteres seguem a ordem Unicode.

Modifique sua variável `pattern` para corresponder a qualquer letra `t` precedida por uma letra minúscula na variável `quote`.Use a faixa de caracteres de `a` a `z` para isso.

# --hints--

Você deve modificar sua variável `pattern` para corresponder a qualquer letra `t` em `quote` precedida por uma letra minúscula usando a classe `[a-z]`.

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "[a-z]t") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0123456789]')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'w[ha]'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
