---
id: 6564c67db26c417561ab510d
title: Passo 31
challengeType: 20
dashedName: step-31
---

# --description--

Como você pode ver na saída, agora seu regex corresponde ao primeiro `l` dentro da string.

No seu padrão, você pode adicionar um quantificador após um caractere para especificar quantas vezes esse caractere deve ser repetido. Por exemplo, o quantificador `+` significa que deve se repetir uma ou mais vezes.

Adicione um quantificador `+` ao seu padrão.

# --hints--

Você deve modificar sua variável `pattern` para `re.compile('l+')`.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('l+')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('l')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
