---
id: 65648f4b2281ba50051ae39c
title: Passo 17
challengeType: 20
dashedName: step-17
---

# --description--

Em seguida, escreva um loop `for` com `i` como a variável do loop.Use a função `range()` para iterar até o valor de `length`.

Dentro do loop, use o operador de atribuição de adição para adicionar um caractere aleatório de `all_characters` ao valor atual de `password`.Use a função `choice()` do módulo `secrets` para isso.

# --hints--

Você deve escrever um loop `for` que itere sobre `range(length)`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;    
    assert(function_body.match(/for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:/));
  }
})
```

Você deve usar o operador `+=` para adicionar um caractere aleatório de `all_characters` ao valor atual de `password`.

```js
({ test: () =>
  {
    const generate_pwd = __helpers.python.getDef(code, "generate_password");
    const {function_body} = generate_pwd;    
    assert(function_body.match(/^(\s*)for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:\s*^\1\1password\s*\+=\s*secrets\.choice\s*\(\s*all_characters\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--    
    password = ''
    # Generate password
    
--fcc-editable-region--
```
