---
id: 6564683821b2ee3174e7250e
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Você pode acessar as utilidades definidas dentro do módulo importado através da notação de ponto. A notação de ponto funciona adicionando um ponto seguido pelo nome da utilidade ao nome do módulo. Por exemplo, aqui está como acessar a constante `ascii_lowercase`:

```py
import string


print(string.ascii_lowercase)
# Output: abcdefghijklmnopqrstuvwxyz
```

Neste projeto, você vai usar diferentes constantes do módulo `string`. 

Declare uma nova variável chamada `letters` e atribua `string.ascii_letters` a essa variável.

# --hints--

Você deve declarar uma variável chamada `letters`.

```js
({ test: () => assert(__userGlobals.has("letters")) })
```

Você deve atribuir `string.ascii_letters` à sua variável `letters`.

```js
({ test: () => assert(runPython(`
    import string
    letters == string.ascii_letters
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


--fcc-editable-region--
```
