---
id: 656475bbf1c2573de1d2c69c
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Toda vez que o código for executado, você deve ver um caractere aleatório da string `all_characters`.Isso é exatamente o que você quer alcançar para criar uma senha aleatória.

No entanto, o algoritmo no qual `random` depende torna os números pseudoaleatórios gerados previsíveis.Portanto, embora o módulo `random` seja adequado para as aplicações mais comuns, ele não pode ser usado para fins criptográficos devido à sua natureza determinística. 

Em vez de importar `random`, importe o módulo `secrets`.Então altere a chamada `print()` para usar `secrets.choice(all_characters)`.

# --hints--

Você deve importar o módulo `secrets` em vez do módulo `random`.

```js
({
  test: () => {
    assert.match(code, /^import\s+secrets/m)
    assert.isFalse( /^import\s+random/m.test(code))
  }
})
```

Você deve modificar sua chamada existente `print()` para imprimir `secrets.choice(all_characters)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*secrets\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
print(random.choice(all_characters))
--fcc-editable-region--
```
