---
id: 6564f0279e23ce924eedd1b2
title: Passo 51
challengeType: 20
dashedName: step-51
---

# --description--

Como o caractere underscore é um caractere válido para nomes de variáveis, ele está incluído na classe de caracteres `\w` (equivalente a `[a-zA-Z0-9_]`), que pode ser usada convenientemente para corresponder a nomes de variáveis.

Portanto, a classe de caracteres `\W` é equivalente a `[^a-zA-Z0-9_]` com o caractere underscore que não é correspondido.Por essa razão você não pode usá-lo para corresponder a todos os seus caracteres especiais.

Delete as últimas três linhas no seu código.

# --hints--

Você deve deletar as últimas três linhas no seu código.

```js
({ test: () => {
  const commentless_code = __helpers.python.removeComments(code);
  assert.notMatch(commentless_code, /pattern\s*=\s*r("|')\\W\1/);
  assert.notMatch(commentless_code, /quote\s*=\s*("|')_\1/);
  assert.notMatch(commentless_code, /print\(\s*re\.findall\(\s*pattern\s*,\s*quote\s*\)\s*\)/);
}})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'\W')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = r'\W'
quote = '_'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
