---
id: 6564642ce403fe29cc23fee7
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Um módulo Python é um arquivo contendo código projetado para executar tarefas específicas.A biblioteca padrão do Python contém muitos módulos que você pode importar e usar no seu código.Você pode conseguir isso usando a declaração `import` seguida do nome do módulo.

Comece este projeto importando o módulo `string`.

# --hints--

Você deve importar o módulo `string`.

```js
({ test: () => assert.match(code, /^import\s+string/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
