---
id: 656485a8a3496d4a36b1496a
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Essas três variáveis constituem os caracteres possíveis para escolher ao gerar a senha.

Logo antes deles, adicione um comentário dizendo `Define the possible characters for the password`.

# --hints--

Você deve adicionar o comentário logo acima da variável `letters`.

```js
({ test: () => assert.match(code, /^#\s*Define the possible characters for the password.*^letters/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string



letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--
```
