---
id: 6564a767a7241362f7d8d664
title: Passo 18
challengeType: 20
dashedName: step-18
---

# --description--

Um underscore único isolado é usado para representar um valor com o qual você não se importa ou que não será usado no seu código. Sua variável de iteração na verdade não é usada.

Modifique sua variável `i` para um único underscore.

# --hints--

Você deve modificar sua variável `i` para `_`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;     
    assert(function_body.match(/for\s+_\s+in\s+range\s*\(\s*length\s*\)\s*:/));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--
    password = ''
    # Generate password
    for i in range(length):
        password += secrets.choice(all_characters)
    
--fcc-editable-region--
```
