---
id: 65646ffeaed2d238c562a014
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

É uma convenção comum colocar as declarações `import` no topo do seu código.E adicionalmente, no caso de múltiplas declarações `import`, ordene-as em ordem alfabética para melhorar a legibilidade.

No início do seu código, importe o módulo `random`.

# --hints--

Você deve importar o módulo `random`.

```js
({ test: () => assert.match(code, /^import\s+random/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string
--fcc-editable-region--

# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)

```
