---
id: 6564d75a923d21815caaa445
title: Passo 43
challengeType: 20
dashedName: step-43
---

# --description--

Se você precisar corresponder a um caractere que tenha um significado especial no padrão, como `.` ou `+`, você pode escapá-lo adicionando um caractere de barra invertida, `\`. Por exemplo, isso corresponde a um sinal de mais literal: `\+`.

Modifique `pattern` para que ele corresponda a um único ponto literal.

# --hints--

Você deve modificar sua variável `pattern` para `\.`.

```js
({ test: () => assert.match(code, /^pattern\s*=\s*("|')\\\.\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '.+'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
