---
id: 65649a631440b757f9f7b428
title: Passo 28
challengeType: 20
dashedName: step-28
---

# --description--

Uma expressão regular, ou regex, é um padrão usado para corresponder a uma combinação específica de caracteres dentro de uma string. É uma alternativa válida às declarações condicionais `if`/`else` quando você precisa corresponder ou encontrar padrões dentro de uma string para fins de validação, substituição de caracteres e outros.

A função `compile()` do módulo `re` compila a string passada como argumento em um objeto de expressão regular que pode ser usado por outros métodos `re`.

Declare uma nova variável `pattern` e atribua o valor de `re.compile('i')` a essa variável.

# --hints--

Você deve declarar uma variável `pattern`.

```js
({ test: () => assert(__userGlobals.has("pattern")) })
```

Você deve atribuir `re.compile('i')` à sua variável `pattern`.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('i')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--

--fcc-editable-region--
```
