---
id: 656ef89b4b486b7a16077864
title: Passo 26
challengeType: 20
dashedName: step-26
---

# --description--

Uma tupla é outra estrutura de dados embutida em Python. Tuplas são muito parecidas com listas, mas são definidas com parênteses `()`, em vez de colchetes. Além disso, tuplas são imutáveis, ao contrário de listas.

```py
my_tuple = ('larch', 1, True)
```

Sua lista de `constraints` vai armazenar tuplas. O primeiro item de cada tupla será um parâmetro de restrição.

Modifique a atribuição da lista `constraints` adicionando uma tupla à sua lista. Use `nums` como o primeiro item e uma string vazia como o segundo item.

# --hints--

Você deve adicionar uma tupla à lista `constraints`.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(.*\)\s*\]/) })
```

Você deve adicionar `(nums, '')` à lista `constraints`.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*("|')\1\s*\)\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = []
        
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
