---
id: 656ef5bc5c0cd464be1df675
title: Passo 23
challengeType: 20
dashedName: step-23
---

# --description--

Em seguida, você vai dar à sua função mais parâmetros que agirão como restrições para a senha gerada.

Modifique a declaração da sua função adicionando `nums`, `special_chars`, `uppercase` e `lowercase` nesta ordem após o parâmetro `length` existente.

# --hints--

Sua função deve receber `length`, `nums`, `special_chars`, `uppercase` e `lowercase` como parâmetros.A ordem importa.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(generate_password))
    sig == '(length, nums, special_chars, uppercase, lowercase)'
  `))
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

--fcc-editable-region--
def generate_password(length):
--fcc-editable-region--    
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password

# new_password = generate_password(8)
# print(new_password)
```
