---
id: 656ef54070c72161e6feb90f
title: Passo 22
challengeType: 20
dashedName: step-22
---

# --description--

Parece tudo bem, mas seria bom ter uma forma de verificar se a senha gerada cumpre características específicas.Por exemplo, um número mínimo de caracteres especiais, dígitos ou letras maiúsculas/minúsculas.Você vai cuidar disso muito em breve.

Por enquanto, comente as duas últimas linhas do seu código.

# --hints--

Você deve transformar as duas últimas linhas do seu código em comentários.

```js
({ test: () => {
  assert.match(code, /#\s*new_password\s*=\s*generate_password\s*\(\s*8\s*\)/);
  assert.match(code, /#\s*print\s*\(\s*new_password\s*\)/);
} })
```


# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--
new_password = generate_password(8)
print(new_password)
--fcc-editable-region--
```
