---
id: 656eec2f40d18056cc58b229
title: Passo 30
challengeType: 20
dashedName: step-30
---

# --description--

O valor `None` é retornado já que `'i'` não é encontrado dentro da string analisada.

Agora, modifique a string passada para `re.compile()` para `'l'` e veja o resultado.

# --hints--

Você deve modificar sua variável `pattern` para `re.compile('l')`.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('l')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)    
--fcc-editable-region--
pattern = re.compile('i')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
