---
id: 656f0d6bc60b58b9777a36c9
title: Passo 34
challengeType: 20
dashedName: step-34
---

# --description--

Uma classe de caractere é indicada por colchetes e corresponde a um caractere entre os especificados entre os colchetes. Por exemplo, `ma[dnt]` corresponde a `mad`, `man` ou `mat`.

Modifique seu padrão para corresponder a um `w` seguido de `h` ou `a`.

# --hints--

Seu padrão deve usar uma classe de caracteres para corresponder a um `w` seguido de `h` ou `a`.

```js
({ test: () => assert(runPython(`
    pattern == "w[ha]" or pattern == "w[ah]"    
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'l+'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
