---
id: 6577333feab1e8e927014f03
title: Passo 63
challengeType: 20
dashedName: step-63
---

# --description--

Ter `all([expression for i in iterable])` significa que uma nova lista é criada avaliando `expression` para cada `i` em `iterable`.Após a função `all()` iterar sobre a lista recém-criada, a lista é deletada automaticamente, pois não é mais necessária.

Memória pode ser economizada usando uma expressão geradora.Expressões geradoras seguem a sintaxe de compreensões de listas mas usam parênteses em vez de colchetes.

Transforme sua list comprehension em uma generator expression removendo os colchetes.

# --hints--

Você deve transformar sua list comprehension em uma generator expression removendo os colchetes.

```js
({ test: () => assert.match(code, /^(\s*)if\s+all\s*\(\s*constraint\s*<=\s*len\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*password\s*\)\s*\)\s*for\s+constraint\s*,\s*pattern\s+in\s+constraints\s*\)\s*:\s*^\1\s{4}break/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        if all(
            [
                constraint <= len(re.findall(pattern, password))
                for constraint, pattern in constraints
            ]
        ):
--fcc-editable-region--
            break
    
    return password

# new_password = generate_password(8)
# print(new_password)
```
