---
id: 6577319039f4f7de9251b822
title: Passo 61
challengeType: 20
dashedName: step-61
---

# --description--

Em vez de usar um loop e uma variável contador, você pode alcançar o mesmo resultado com uma abordagem diferente, que você vai implementar nos próximos passos.

`all()` é uma função embutida do Python que retorna `True` se todos os elementos dentro de um iterável dado avaliarem para `True`. Caso contrário, retorna `False`.

Substitua seu `for` loop existente e duas declarações `if` por uma única declaração `if`. Para a condição `if`, use uma chamada para a função `all()` e passe uma lista vazia como argumento para a chamada da função.

# --hints--

Você deve substituir seu `for` loop e duas declarações `if` por uma única declaração `if`. 

```js
({ test: () => {
  assert.match(code, /^(\s{8})if\s+.+:\s*^\1\s{4}break/ms);
  assert.isFalse(/if\s+count\s*==\s*4\s*:/.test(code));
} })
```

Sua nova condição `if` deve ser `all([])`.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])") != _Node()
`)) })
```

Você deve ter um `break` dentro do corpo do seu novo `if`.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])").find_bodies()[0].is_equivalent("break")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        for constraint, pattern in constraints:
            if constraint <= len(re.findall(pattern, password)):
                count += 1
            
        if count == 4:
--fcc-editable-region--
            break

    return password

# new_password = generate_password(8)
# print(new_password)
```
