---
id: 657732654845d8e2fb1217e6
title: Passo 62
challengeType: 20
dashedName: step-62
---

# --description--

No momento, `all()` está recebendo uma lista vazia como argumento.Preencha essa lista vazia usando a sintaxe de compreensão para que a lista armazene os resultados da avaliação da expressão `constraint <= len(re.findall(pattern, password))` para cada tupla `constraint`-`pattern` na lista `constraints`.

Dessa forma, você sairá do loop `while` somente depois que todos os requisitos forem cumpridos.

# --hints--

Você deve passar `[constraint <= len(re.findall(pattern, password)) for constraint, pattern in constraints]` para a função `all()`.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("all([constraint <= len(re.findall(pattern, password)) for constraint, pattern in constraints])")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--        
        if all([]):
--fcc-editable-region--        
            break
    
    return password

# new_password = generate_password(8)
# print(new_password)
```
