---
id: 657db4cb77190e33a20e852a
title: Passo 69
challengeType: 20
dashedName: step-69
---

# --description--

Quando você combina argumentos padrão com argumentos nomeados, é possível passar explicitamente menos argumentos do que os exigidos pela função.Os argumentos que não são explicitamente passados para a chamada da função irão assumir seus valores padrão.

Modifique sua chamada `generate_password()` para aceitar apenas `length=8`.

# --hints--

Você deve passar apenas `length=8` para sua chamada `generate_password()`.

```js
({ test: () => assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*length\s*=\s*8\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length=16, nums=1, special_chars=1, uppercase=1, lowercase=1):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (special_chars, fr'[{symbols}]'),
            (uppercase, r'[A-Z]'),
            (lowercase, r'[a-z]')
        ]

        # Check constraints        
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):
            break
    
    return password
    
--fcc-editable-region--
new_password = generate_password(nums=1, length=8, special_chars=1, uppercase=1, lowercase=1)
print(new_password)
--fcc-editable-region--
```
