---
id: 68af273960e009597dd861e2
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

デバッグや出力のために読みやすい表示が欲しいです。`__str__` メソッドを作成して、`title (year) - duration min, director` の形式の文字列を返してください（`title`、`year`、`duration`、`director` はそれぞれ対応する属性の値に置き換えられるべきです）。これにより、映画を出力するとわかりやすい概要が表示されます。

# --hints--

`Movie` クラスには `__str__` メソッドが必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__str__")`)
})
```

`__str__` メソッドは `title (year) - duration min, director` の形式の文字列を返すべきで、`title`、`year`、`duration`、`director` はそれぞれ対応する属性の値に置き換えられるべきです。

```js
({
    test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    print(movie1)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "The Matrix (1999) - 136 min, The Wachowskis" in output
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    
--fcc-editable-region--
```
