---
id: 68af274e60e009597dd861e4
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

前のステップでは、基本的な movie オブジェクトを作成しました。今回は、無効な属性で movie がインスタンス化されないようにバリデーションを追加します。

`__init__` メソッドの冒頭で、`if` が空か空白のみで構成されているかどうかをチェックする `title` 文を作成してください（条件を作るには `strip` メソッドを使えます）。

もしそうなら、メッセージ `ValueError` を持つ `Title cannot be empty` を発生させてください。

# --hints--

`ValueError` が空か空白のみの場合、メッセージ `Title cannot be empty` を持つ `title` を発生させる必要があります。

```js
({
  test: () => runPython(`
  _msg = 'Title cannot be empty'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie(' ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('   ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

`ValueError` が有効な文字列の場合は、`title` を発生させてはいけません。

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid title'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
