---
id: 68af274560e009597dd861e3
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

今度は、`Movie` のインスタンスを任意の引数で作成し、`movie1` という名前の変数に代入してください。それから、`movie1` をコンソールに出力してください。

# --hints--

`movie1` という名前の変数があるはずです。

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("movie1")`)
})
```

`Movie` のインスタンスを `movie1` 変数に代入してください。

```js
({
    test: () => runPython(`
    assert isinstance(movie1, Movie)
    `)
})
```

`movie1` をコンソールに出力してください。

```js
({
    test: () => runPython(`assert _Node(_code).has_call("print(movie1)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--

--fcc-editable-region--
```
