---
id: 68af277e60e009597dd861e9
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

新しいクラス`MediaCatalogue`を作成し、`__init__`メソッドを追加してください。このメソッドは`items`属性を空のリストに設定します。これはカタログ内のアイテムのリストを表します。

# --hints--

`MediaCatalogue`という名前のクラスがあるはずです。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("MediaCatalogue")`)
})
```

`MediaCatalogue`クラスには`__init__`メソッドが必要です。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_class("MediaCatalogue").has_function("__init__")
    `)
})
```

`__init__`メソッドは`items`属性を空のリストに設定してください。

```js
({
    test: () => runPython(`
    _catalogue = MediaCatalogue()
    assert _catalogue.items == []
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--

--fcc-editable-region--
try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')
```
