---
id: 68af275860e009597dd861e5
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

最初の映画は1895年に作られました。`if`が`year`より小さいかどうかをチェックする別の`1895`文を作成し、メッセージ`ValueError`で`Year must be 1895 or later`を発生させてください。

# --hints--

`ValueError`が`Year must be 1895 or later`より小さい場合、メッセージ`year`で`1895`を発生させる必要があります。

```js
({
  test: () => runPython(`
  _msg = 'Year must be 1895 or later'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie('The Matrix', 1894, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('The Matrix', 0, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

`ValueError`が有効な数字の場合は、`year`を発生させてはいけません。

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1895, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid year'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
