---
id: 68af278a60e009597dd861ea
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

`MediaCatalogue` クラスには、メディアアイテムをカタログに追加するメソッドが必要です。アイテム `add` を受け取り、`media_item` リストに追加する `items` メソッドを作成してください。

# --hints--

`MediaCatalogue` には `add` メソッドが必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").has_function("add")`)
})
```

`add` メソッドは、2つのパラメータ `self` と `media_item` を持つ必要があります。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("add").has_args("self, media_item")`)
})
```

`add` メソッドは、渡されたアイテムを `items` 属性に追加する必要があります。

```js
({
    test: () => runPython(`
    _catalogue = MediaCatalogue()
    _catalogue.add("first")
    assert _catalogue.items == ["first"]
    _catalogue.add("second")
    assert _catalogue.items == ["first", "second"]
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    
--fcc-editable-region--
try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')
```
