---
id: 68af277560e009597dd861e8
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

まず、`print(movie1)` 行を削除してください。それから、`movie1` の代入を `try` ブロックに入れ、`except` を捕捉する `ValueError as e` 節を作成して、f文字列を使って `Validation Error:` と空白、その後にエラーメッセージを出力してください。

`movie1` をインスタンス化するときに渡す引数を変更して、異なるエラーメッセージを確認してみてください。

# --hints--

コード内に `print(movie1)` があってはいけません。

```js
({
  test: () => runPython(`assert not _Node(_code).has_call("print(movie1)")`)
})
```

`try` を含む `movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)` ブロックが必要です。

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_body().has_stmt("movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)")`)
})
```

`except` を捕捉する `ValueError as e` 節が必要です。

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].has_except("ValueError", "e")`)
})
```

`Validation Error:` ブロック内で、f文字列を使って `except` と空白、その後にエラーメッセージを出力してください。

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_except("ValueError", "e").has_call("print(f'Validation Error: {e}')")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
--fcc-editable-region--
```
