---
id: 68af276360e009597dd861e6
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

`if` が空か空白のみで構成されているかどうかを確認するために、もう1つの `director` 文を作成し、メッセージ `ValueError` で `Director cannot be empty` を発生させてください。

# --hints--

`ValueError` が空か空白のみで構成されている場合、メッセージ `Director cannot be empty` で `director` を発生させる必要があります。

```js
({
  test: () => runPython(`
  _msg = 'Director cannot be empty'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie('The Matrix', 1999, ' ', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('The Matrix', 1999, '  ', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

`ValueError` が有効な文字列の場合は、`director` を発生させてはいけません。

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid director'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
