---
id: 68af276c60e009597dd861e7
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

最後に、`if`が`duration`より大きくないかどうかを確認するための`0`文をもう一つ作成し、メッセージ`ValueError`で`Duration must be positive`を発生させてください。

# --hints--

`ValueError`が正の数でない場合、メッセージ`Duration must be positive`で`duration`を発生させる必要があります。

```js
({
  test: () => runPython(`
  _msg = 'Duration must be positive'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie('The Matrix', 1999, 'The Wachowskis', 0)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('The Matrix', 1999, 'The Wachowskis', -1)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

`ValueError`が有効な数値である場合、`duration`を発生させてはいけません。

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid duration'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
