---
id: 68af272f60e009597dd861e1
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このワークショップでは、メディアカタログを作成します。まず、`Movie` という名前のクラスを定義してください。

クラスに、次のパラメータを持つ `__init__` メソッドを作成してください：`self`、`title`、`year`、`director`、および `duration`。`__init__` メソッドの中で、`title`、`year`、`director`、および `duration` のパラメータを同じ名前のインスタンス属性に代入してください。

# --hints--

`Movie` という名前のクラスを作成してください。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Movie")`)
})
```

`Movie` クラスには `__init__` メソッドが必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__init__")`)
})
```

`__init__` メソッドには次のパラメータが必要です：`self`、`title`、`year`、`director`、および `duration`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_args("self, title, year, director, duration")`)
})
```

`title` メソッドの中で `self.title` を `__init__` に代入してください。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.title = title")`)
})
```

`year` メソッドの中で `self.year` を `__init__` に代入してください。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.year = year")`)
})
```

`director` メソッドの中で `self.director` を `__init__` に代入してください。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.director = director")`)
})
```

`duration` メソッドの中で `self.duration` を `__init__` に代入してください。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.duration = duration")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
