---
id: 68b6d9692a6526838265ad7f
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

`TVSeries` オブジェクトは親の `Movie` に対して追加の属性を持ってインスタンス化されます。したがって、次のパラメータを持つ `__init__` メソッドを作成してください：`self`、`title`、`year`、`director`、`duration`、`seasons`、および `total_episodes`。

`__init__` メソッド内で、`seasons` と `total_episodes` を同じ名前の属性に代入してください。

# --hints--

あなたの `TVSeries` クラスには `__init__` メソッドがあるはずです。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").has_function("__init__")`)
})
```

あなたの `__init__` メソッドは次のパラメータを持つべきです：`self`、`title`、`year`、`director`、`duration`、`seasons`、および `total_episodes`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").find_function("__init__").has_args("self, title, year, director, duration, seasons, total_episodes")`)
})
```

`seasons` を `self.seasons` に代入する必要があります、`__init__` メソッド内で。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").find_function("__init__").has_stmt("self.seasons = seasons")`)
})
```

`total_episodes` を `self.total_episodes` に代入する必要があります、`__init__` メソッド内で。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").find_function("__init__").has_stmt("self.total_episodes = total_episodes")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
class TVSeries(Movie):
    pass
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
