---
id: 68b704862dd6559e890c2a86
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

`if` 文の後に、`result` という名前の変数を作成してください。これはメソッドから返す文字列になります。

そのため、`Media Catalogue (N items):` の形式の f-string を代入してください。`N` を `items` リストのアイテム数に置き換え、文字列の末尾に改行コードを2つ付けてください。

# --hints--

`result` メソッド内に `__str__` という名前の変数があるはずです。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("__str__").has_variable("result")`)
})
```

`f'Media Catalogue ({len(self.items)} items):\n\n'` 変数に `result` を代入してください。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("__str__").has_stmt("result = f'Media Catalogue ({len(self.items)} items):\\\\n\\\\n'")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)
--fcc-editable-region--
    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'
        
--fcc-editable-region--
try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')
```
