---
id: 68b813ecdbe8896fd1fcf736
title: ステップ 23
challengeType: 20
dashedName: step-23
---

# --description--

`add` の `catalogue` メソッドを使って、`series1` リストに `items` を追加してください。

次に、別の変数 `series2` を作成します。`TVSeries` の別のインスタンスを代入し、`series2` もカタログに追加してください。

# --hints--

`add` の `catalogue` メソッドを使い、`series1` を渡してください。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_trys()[0].find_body().has_stmt("catalogue.add(series1)")
    `)
})
```

`series2` という名前の変数を作成し、`TVSeries` のインスタンスを代入してください。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_trys()[0].find_body().has_variable("series2")
    assert isinstance(series2, TVSeries)
    `)
})
```

`add` の `catalogue` メソッドを使い、`series2` を渡してください。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_trys()[0].find_body().has_stmt("catalogue.add(series2)")
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class TVSeries(Movie):
    def __init__(self, title, year, director, duration, seasons, total_episodes):
        super().__init__(title, year, director, duration)

        if seasons < 1:
            raise ValueError('Seasons must be 1 or greater')
        if total_episodes < 1:
            raise ValueError('Total episodes must be 1 or greater')
        
        self.seasons = seasons
        self.total_episodes = total_episodes

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.seasons} seasons, {self.total_episodes} episodes, {self.duration} min avg, {self.director}'

class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
--fcc-editable-region--
    series1 = TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 182)

--fcc-editable-region--
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
