---
id: 68b804f3be853a355f055fef
title: ステップ 18
challengeType: 20
dashedName: step-18
---

# --description--

`super` 関数は、親クラスを明示的に名前を指定せずに参照するために使います。例えば、`super().__init__(arg1, arg2, ...)` は親クラスで定義された `__init__` メソッドを呼び出します。

現在、`__init__` メソッドを `Movie` から継承したものを実質的にオーバーライドしています。つまり、`TVSeries` オブジェクトは `seasons` 属性と `total_episodes` 属性だけを持つことになります。これは望んでいる状態ではありません。

`Movie` クラスの `__init__` に定義されたすべてのコードを保持するには、親の `super` を呼び出すために `__init__` を使ってください。メソッドに必要な引数を渡す必要があることに注意してください。この場合は `title`、`year`、`director`、および `duration` です。

# --hints--

`super().__init__(title, year, director, duration)` を `__init__` メソッド内で呼び出す必要があります。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").find_function("__init__").has_stmt("super().__init__(title, year, director, duration)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
class TVSeries(Movie):
    def __init__(self, title, year, director, duration, seasons, total_episodes):
        
        self.seasons = seasons
        self.total_episodes = total_episodes
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
