---
id: a6e40f1041b06c996f7b2406
title: 最初の要素を見つける関数を作成する
challengeType: 26
dashedName: build-a-first-element-finder
---

# --description--

このラボでは、配列を調べてテスト関数（「真偽テスト」とも呼ばれます）を通過する最初の要素を返す関数を作成します。

関数は配列を繰り返し処理し、提供されたテスト関数を使って各要素をテストします。最後に、テスト関数が`true`を返す最初の要素を返します。

例：

```js
findElement([1, 3, 5, 8], num => num % 2 === 0) // returns 8
findElement([1, 3, 5], num => num % 2 === 0)    // returns undefined
```

**目的：** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー：**

1. `findElement`という名前の関数があり、配列と関数を引数として受け取る必要があります。
2. 関数は真偽テストを通過する配列内の最初のアイテムを返す必要があります。つまり、渡された関数`func`を、要素`x`に対して呼び出したとき、`func(x)`が`true`であれば真偽テストを通過したことになります。
3. どの要素もテストを通過しなかった場合、関数は`undefined`を返す必要があります。

# --hints--

`findElement`関数が必要です。

```js
assert.isFunction(findElement);
```

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })`は`8`を返すべきです。

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })`は`undefined`を返すべきです。

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

`findElement([1, 2, 3, 4], function(num) { return num > 2; })`は`3`を返すべきです。

```js
assert.strictEqual(
  findElement([1, 2, 3, 4], function (num) {
    return num > 2;
  }),
  3
);
```

`findElement(["hello", "world", "javascript"], function(str) { return str.length > 5; })`は`"javascript"`を返すべきです。

```js
assert.strictEqual(
  findElement(["hello", "world", "javascript"], function (str) {
    return str.length > 5;
  }),
  "javascript"
);
```

`findElement(["cat", "dog", "bird"], function(str) { return str.length > 10; })`は`undefined`を返すべきです。

```js
assert.strictEqual(
  findElement(["cat", "dog", "bird"], function (str) {
    return str.length > 10;
  }),
  undefined
);
```

`findElement([2, 4, 6, 8], function(num) { return num % 2 === 0; })`は`2`を返すべきです。

```js
assert.strictEqual(
  findElement([2, 4, 6, 8], function (num) {
    return num % 2 === 0;
  }),
  2
);
```

`findElement([], function(num) { return num > 0; })`は`undefined`を返すべきです。

```js
assert.strictEqual(
  findElement([], function (num) {
    return num > 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
