---
id: 656469e62377ab34a5bcb8d1
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Declare two new variables named `digits` and `symbols` and assign them `string.digits` and `string.punctuation`, respectively.

# --hints--

Du solltest eine Variable mit dem Namen `digits` deklarieren.

```js
({ test: () => assert(__userGlobals.has("digits")) })
```

Du sollltest deiner `digits`-Variablen `string.digits` zuweisen.

```js
({ test: () => assert(runPython(`
    import string
    digits == string.digits
  `))
})
```

Du solltest eine Variable mit dem Namen `symbols` deklarieren.

```js
({ test: () => assert(__userGlobals.has("symbols")) })
```

Du sollltest deiner `symbols`-Variablen `string.punctuation` zuweisen.

```js
({ test: () => assert(runPython(`
    import string
    symbols == string.punctuation
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string

letters = string.ascii_letters
--fcc-editable-region--
```
