---
id: 6564683821b2ee3174e7250e
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

You can access the utilities defined inside the imported module through the dot notation. The dot notation works by appending a dot followed by the utility name to the module name. For example, here's how to access the `ascii_lowercase` constant:

```py
import string


print(string.ascii_lowercase)
# Output: abcdefghijklmnopqrstuvwxyz
```

In this project, you are going to use different constants from the `string` module.

Declare a new variable called `letters` and assign `string.ascii_letters` to this variable.

# --hints--

Du solltest eine Variable mit dem Namen `letters` deklarieren.

```js
({ test: () => assert(__userGlobals.has("letters")) })
```

Du sollltest deiner `letters`-Variablen `string.ascii_letters` zuweisen.

```js
({ test: () => assert(runPython(`
    import string
    letters == string.ascii_letters
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


--fcc-editable-region--
```
