---
id: 587d78af367417b2b2512b03
title: 설문조사 양식 만들기
challengeType: 14
saveSubmissionToDB: true
forumTopicId: 301145
dashedName: build-a-survey-form
---

# --description--

**목표:** <a href="https://survey-form.freecodecamp.rocks" target="_blank" rel="noopener noreferrer nofollow">https://survey-form.freecodecamp.rocks</a>와 기능적으로 유사한 앱을 만드세요. **이 데모 프로젝트를 복사하지 마세요**.

**사용자 스토리:**

1. `h1` 요소에 `id`가 `title`인 페이지 제목이 있어야 합니다.
1. `p` 요소에 `id`가 `description`인 간단한 설명이 있어야 합니다.
1. `form` 요소에 `id`가 `survey-form`인 요소가 있어야 합니다.
1. form 요소 안에 `input` 필드에 이름을 입력해야 하며, 이 필드는 `id`가 `name`이고 `type`가 `text`이어야 합니다.
1. form 요소 안에 `input` 필드에 이메일을 입력해야 하며, 이 필드는 `id`가 `email`이어야 합니다.
1. 올바른 형식이 아닌 이메일을 입력하면 HTML5 유효성 검사 오류가 표시됩니다.
1. form 안에 `input` 필드에 숫자를 입력할 수 있어야 하며, 이 필드는 `id`가 `number`이어야 합니다.
1. 숫자 입력 필드는 숫자가 아닌 값을 입력하지 못하게 하거나(브라우저에 따라 다름) HTML5 유효성 검사 오류를 보여야 합니다.
1. `min`와 `max` 속성으로 정의된 숫자 입력 범위를 벗어난 숫자를 입력하면 HTML5 유효성 검사 오류가 표시됩니다.
1. 이름, 이메일, 숫자 입력 필드 각각에 대해 form 안에 `label` 요소가 있어야 하며, 각각 `id="name-label"`, `id="email-label"`, `id="number-label"` 아이디로 각 필드의 목적을 설명해야 합니다.
1. 이름, 이메일, 숫자 입력 필드 각각에 대해 설명이나 지침을 제공하는 자리 표시자 텍스트가 있어야 합니다.
1. form 요소 안에 `select` 드롭다운 요소가 있어야 하며, `id`가 `dropdown`이고 선택할 수 있는 옵션이 최소 두 개 이상 있어야 합니다.
1. form 요소 안에 `name` 속성으로 그룹화된 최소 두 개 이상의 라디오 버튼 그룹에서 옵션을 선택할 수 있어야 합니다.
1. form 요소 안에 `value` 속성이 있는 여러 체크박스 중에서 여러 개를 선택할 수 있어야 합니다.
1. form 요소 안에 추가 의견을 위한 `textarea`가 있어야 합니다.
1. form 요소 안에 모든 입력을 제출할 `id`가 `submit`인 버튼이 있어야 합니다.

사용자 스토리를 모두 충족하고 아래 모든 테스트를 통과해야 이 프로젝트를 완료할 수 있습니다. 자신만의 스타일로 만드세요. 코딩을 즐기세요!

**참고:** HTML에 `<link rel="stylesheet" href="styles.css">`를 추가하여 스타일시트를 연결하고 CSS를 적용하세요.

# --hints--

`h1` 요소에 `id`가 `title`인 요소가 있어야 합니다.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

`#title`는 비어 있으면 안 됩니다.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`p` 요소에 `id`가 `description`인 요소가 있어야 합니다.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

`#description`는 비어 있으면 안 됩니다.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`form` 요소에 `id`가 `survey-form`인 요소가 있어야 합니다.

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

`input` 요소에 `id`가 `name`인 요소가 있어야 합니다.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

`#name`는 `type`가 `text`이어야 합니다.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

`#name`는 입력을 필수로 해야 합니다.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

`#name`는 `#survey-form`의 자손 요소여야 합니다.

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

`input` 요소에 `id`가 `email`인 요소가 있어야 합니다.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

`#email`는 `type`가 `email`이어야 합니다.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

`#email`는 입력을 필수로 해야 합니다.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

`#email`는 `#survey-form`의 자손 요소여야 합니다.

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

`input` 요소에 `id`가 `number`인 요소가 있어야 합니다.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

`#number`는 `#survey-form`의 자손 요소여야 합니다.

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

`#number`는 `type`가 `number`이어야 합니다.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

`#number`는 숫자 값을 가진 `min` 속성이 있어야 합니다.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

`#number`는 숫자 값을 가진 `max` 속성이 있어야 합니다.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

`label` 요소에 `id`가 `name-label`인 요소가 있어야 합니다.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

`label` 요소에 `id`가 `email-label`인 요소가 있어야 합니다.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

`label` 요소에 `id`가 `number-label`인 요소가 있어야 합니다.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

`#name-label`는 입력을 설명하는 텍스트를 포함해야 합니다.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`#email-label`는 입력을 설명하는 텍스트를 포함해야 합니다.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`#number-label`는 입력을 설명하는 텍스트를 포함해야 합니다.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`#name-label`는 `#survey-form`의 자손 요소여야 합니다.

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

`#email-label`는 `#survey-form`의 자손 요소여야 합니다.

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

`#number-label`는 `#survey-form`의 자손 요소여야 합니다.

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

`#name`는 `placeholder` 속성과 값을 가져야 합니다.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

`#email`는 `placeholder` 속성과 값을 가져야 합니다.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

`#number`는 `placeholder` 속성과 값을 가져야 합니다.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

`select` 필드는 `id`가 `dropdown`이어야 합니다.

```js
const el = document.getElementById('dropdown');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'SELECT');
```

`#dropdown`는 선택 가능한(비활성화되지 않은) `option` 요소가 최소 두 개 이상 있어야 합니다.

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert.isAtLeast(els.length, 2);
```

`#dropdown`는 `#survey-form`의 자손 요소여야 합니다.

```js
const el = document.querySelector('#survey-form #dropdown');
assert.isNotNull(el);
```

라디오 버튼인 `input` 요소가 `type`가 `radio`인 것이 최소 두 개 이상 있어야 합니다.

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

라디오 버튼 두 개 이상이 `#survey-form`의 자손 요소여야 합니다.

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

모든 라디오 버튼은 `value` 속성과 값을 가져야 합니다.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][value=""], input[type="radio"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

모든 라디오 버튼은 `name` 속성과 값을 가져야 합니다.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][name=""], input[type="radio"]:not([name])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

라디오 버튼 그룹마다 최소 2개의 라디오 버튼이 있어야 합니다.

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  });
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert(false);
});

assert.isAbove(groupKeys.length, 0);
```

체크박스인 `input` 요소가 `type`가 `checkbox`인 것이 최소 두 개 이상 있어야 하며, `#survey-form`의 자손 요소여야 합니다.

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert.isAtLeast(els.length, 2);
```

`#survey-form` 안의 모든 체크박스는 `value` 속성과 값을 가져야 합니다.

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll(
  '#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

`textarea` 요소가 `#survey-form`의 자손 요소여야 합니다.

```js
const el = document.querySelector('#survey-form textarea');
assert.isNotNull(el);
```

`input` 또는 `button` 요소에 `id`가 `submit`인 요소가 있어야 합니다.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.isTrue(el.tagName === 'INPUT' || el.tagName === 'BUTTON');
```

`#submit`는 `type`가 `submit`이어야 합니다.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.strictEqual(el.type, 'submit');
```

`#submit`는 `#survey-form`의 자손 요소여야 합니다.

```js
const el = document.querySelector('#survey-form #submit');
assert.isNotNull(el);
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
    <title>Survey Form</title>
  </head>
  <body>
    <h1>Survey Form</h1>
    <p>The card below was built as a sample survey form for freeCodeCamp.</p>
    <main id="main">
      <h1 id="title">Join the Togepi Fan Club!</h1>
      <p id="description">
        Enter your information here to receive updates about club activities,
        our monthly newsletter, and other email communications.
      </p>
      <form id="survey-form" action="#">
        <label for="name" id="name-label"
          ><p>Name:</p>
          <input type="text" id="name" placeholder="e.g. John Smith" required />
        </label>
        <label for="email" id="email-label"
          ><p>Email:</p>
          <input
            type="email"
            id="email"
            placeholder="e.g. john.smith@email.com"
            required
          />
        </label>
        <label for="age" id="number-label"
          ><p>Age<em>(optional)</em>:</p>
          <input
            type="number"
            id="number"
            placeholder="e.g. 19"
            min="18"
            max="99"
          />
        </label>
        <label for="interest" id="interest-label"
          ><p>What are you most interested in?</p>
          <select id="dropdown">
            <option selected disabled hidden></option>
            <option id="battles">Battling</option>
            <option id="breeding">Breeding</option>
            <option id="catching">Completing my Pokedex</option>
            <option id="exploring">Exploring new regions</option>
          </select>
        </label>
        <p>Who is your favourite Pokemon?</p>
        <label for="togepi">
          <input
            id="togepi"
            type="radio"
            name="favorite"
            value="togepi"
          />Togepi!
        </label>
        <label for="pikachu">
          <input
            id="pikachu"
            type="radio"
            name="favorite"
            value="pikachu"
          />Pikachu
        </label>
        <label for="other">
          <input id="other" type="radio" name="favorite" value="other" />Other
        </label>
        <p>Which communications do you want to receive?</p>
        <label for="newsletter">
          <input
            id="newsletter"
            type="checkbox"
            name="communications"
            value="newsletter"
          />Newsletter
        </label>
        <label for="events">
          <input
            id="events"
            type="checkbox"
            name="communications"
            value="events"
          />Event updates
        </label>
        <label for="updates">
          <input
            id="updates"
            type="checkbox"
            name="communications"
            value="updates"
          />Club updates
        </label>
        <p>Any other information you would like to share?</p>
        <textarea id="additional-information" rows="4" cols="50">
You can provide comments, suggestions, or feedback here.</textarea
        >
        <p>
          <em
            >Please note: This form is a proof of concept. Submitting the form
            will not actually transmit your data.</em
          >
        </p>
        <button type="Submit" id="submit">Submit</button>
      </form>
    </main>
  </body>
  <footer>
    <a href="../">Return to Project List</a> |
    <a href="https://www.nhcarrigan.com">Return to HomePage</a>
  </footer>
</html>
```

```css
main {
  text-align: center;
  background-color: #92869c;
  background-blend-mode: lighten;
  max-width: 500px;
  margin: 20px auto;
  border-radius: 50px;
  box-shadow: 10px 10px rgba(0, 0, 0, 0.5);
  color: black;
}
body {
  text-align: center;
  background: #3a3240;
  color: white;
}
input,
textarea,
select,
button {
  background: #3a3240;
  color: white;
}
a {
  color: white;
}
```
