---
id: 69860157f7fce39f9df399a3
title: Algorithmus für das N-Damen-Problem implementieren
challengeType: 26
dashedName: implement-the-n-queens-algorithm
---

# --description--

Das N-Damen-Problem fordert Sie auf, N Damen auf einem N×N-Schachbrett so zu platzieren, dass sich keine zwei Damen gegenseitig angreifen (keine zwei teilen sich eine Reihe, Spalte oder Diagonale).

Zum Beispiel, wenn es ein 4x4-Brett gibt, ist eine gültige Anordnung:

```md
[1, 3, 0, 2]
```

Das bedeutet, dass in Reihe 0 die Dame in Spalte 1 steht; in Reihe 1 die Dame in Spalte 3; in Reihe 2 die Dame in Spalte 0; und in Reihe 3 die Dame in Spalte 2.

Visuell sieht diese Anordnung so aus:

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

Dabei steht `Q` für eine Dame und `.` für ein leeres Feld.

In diesem Labor implementieren Sie den Lösungsalgorithmus für das N-Damen-Problem mit dem Ansatz der Tiefensuche.


**Ziel**: Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `dfsNQueens` haben.
2. Die Funktion sollte genau ein Argument akzeptieren: eine ganze Zahl `n`.
3. Wenn `n` kleiner als `1` ist, sollte die Funktion ein leeres Array (`[]`) zurückgeben.
4. Die Funktion sollte eine Liste von Lösungen zurückgeben; jede Lösung ist selbst eine Liste der Länge `n`, wobei das Element am Index `i` den Spaltenindex (0-basiert) der Dame in der Reihe `i` darstellt.

# --hints--

Sie sollten eine Funktion mit dem Namen `dfsNQueens` haben, die ein Argument annimmt.

```js
assert.isFunction(dfsNQueens);
assert.lengthOf(dfsNQueens, 1);
```

Wenn `n` kleiner als `1` ist, sollte die Funktion ein leeres Array zurückgeben.

```js
assert.sameDeepOrderedMembers(dfsNQueens(0), []);
assert.sameDeepOrderedMembers(dfsNQueens(-1), []);
assert.sameDeepOrderedMembers(dfsNQueens(-5), []);
```

Die Funktion sollte ein Array von Lösungen zurückgeben, wobei jede Lösung ein Array der Länge `n` ist.

```js
const output = dfsNQueens(4);
assert.isArray(output);
output.forEach(solution => {
    assert.isArray(solution);
    assert.lengthOf(solution, 4);
})
```

`dfsNQueens(1)` sollte `[[0]]` zurückgeben.

```js
assert.sameDeepOrderedMembers(dfsNQueens(1), [[0]]);
```

`dfsNQueens(2)` sollte `[]` zurückgeben.

```js
assert.sameDeepOrderedMembers(dfsNQueens(2), []);
```

`dfsNQueens(3)` sollte `[]` zurückgeben.

```js
assert.sameDeepOrderedMembers(dfsNQueens(3), []);
```

`dfsNQueens(4)` sollte `[[1, 3, 0, 2], [2, 0, 3, 1]]` zurückgeben.

```js
assert.sameDeepOrderedMembers(dfsNQueens(4), [[1, 3, 0, 2], [2, 0, 3, 1]]);
```

`dfsNQueens(5)` sollte `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]` zurückgeben.

```js
assert.sameDeepOrderedMembers(dfsNQueens(5), [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]);
```

`dfsNQueens(5).length` sollte `10` sein.

```js
assert.lengthOf(dfsNQueens(5), 10);
```

`dfsNQueens(8).length` sollte `92` sein.

```js
assert.lengthOf(dfsNQueens(8), 92);
```

`dfsNQueens` sollte das korrekte Ergebnis zurückgeben.

```js
assert.lengthOf(dfsNQueens(7), 40);
assert.sameDeepOrderedMembers(dfsNQueens(6), [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfsNQueens(n) {
  if (n < 1) return [];

  const results = [];

  function isSafe(queens, row, col) {
    for (let r = 0; r < queens.length; r++) {
      const c = queens[r];
      if (c === col || Math.abs(row - r) === Math.abs(col - c)) {
        return false;
      }
    }
    return true;
  }

  function dfs(queens) {
    const row = queens.length;
    if (row === n) {
      results.push(queens.slice());
      return;
    }

    for (let col = 0; col < n; col++) {
      if (isSafe(queens, row, col)) {
        queens.push(col);
        dfs(queens);
        queens.pop();
      }
    }
  }

  dfs([]);
  return results;
}
```
