---
id: 69de4308e652aaeb5ccdb03e
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このワークショップでは、Pythonの文字列メソッドを使う練習をします。

`email`という名前の変数を作成し、文字列`alice.johnson@company.com`を代入してください。

# --hints--

`email`という名前の変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email')
`) })
```

`email`に文字列`alice.johnson@company.com`を代入してください。文字列はシングルクォートか二重引用符で囲むことを忘れないでください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('email').is_equivalent("email = 'alice.johnson@company.com'")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
