---
id: 69e20876baed33de636b989a
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

文字列の中で特定の文字の位置を見つけたい場合があります。Pythonはこの目的のために組み込みの `.find()` メソッドを提供しています。これは、指定した部分文字列が最初に現れる位置のインデックス（位置）を表す整数を返します。

例を示します。

```py
text = 'hello'
print(text.find('e'))  # 1
```

次に、`email` に対して `.find()` メソッドを呼び出し、`'@'` を渡して `email` 文字列内のアットマークの位置を見つけてください。その結果を `at_position` という変数に代入してください。

その後、`at_position` を出力して、結果が `Position of @: 13` と表示されるようにしてください。

# --hints--

`at_position` という変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_position')
`) })
```

`at_position` に `email.find('@')` の値を代入してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_position').is_equivalent("at_position = email.find('@')")
`) })
```

`at_position` を出力して、結果が `Position of @: 13` と表示されるようにしてください。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Position of @: {at_position}')",
    "print('Position of @: ' + str(at_position))",
    "print('Position of @:', at_position)",
    "print('Position of @:', str(at_position))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')
--fcc-editable-region--

--fcc-editable-region--
```
