---
id: 69e20a0451c7a3ef189514d3
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

メール内の`@`の位置がわかったので、`@`の前にあるユーザー名を抽出できます。

前のワークショップでは、スライスを使って文字列の一部を抽出しました。復習すると、構文は次の通りです：

```python
text = 'hello world'
print(text[:5])  # hello
```

`email`にスライスを使って`at_position`の前のすべてを抽出し、その結果を`username`という変数にストアしてください。

次に`username`を出力して、結果が`Username: alice.johnson`となるようにしてください。

# --hints--

`username`という変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('username')
`) })
```

`email`に`at_position`を使ったスライスをして、`username`に値を代入してください。

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('username')
assert node.is_equivalent("username = email[:at_position]") or node.is_equivalent("username = email[0:at_position]")
`) })
```

`username`を出力して、結果が`Username: alice.johnson`となるようにしてください。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {username}')",
    "print('Username: ' + username)",
    "print('Username:', username)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

--fcc-editable-region--

--fcc-editable-region--
```
