---
id: 69e6284f3b7659c5a05c2a79
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

`in` キーワードを使って、文字列の中に部分文字列が存在するかどうかを確認できます。部分文字列が存在すれば `True` を返し、そうでなければ `False` を返します。例を示します：

```py
sentence = 'I love Python'
print('Python' in sentence) # True
print('Java' in sentence) # False
```

`email` に文字 `@` が含まれているかどうかをチェックする `has_at` という変数を作成してください。

次に `has_at` を出力して、結果が `Has @ symbol: True` と表示されるようにしてください。

# --hints--

`has_at` という名前の変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_at')
`) })
```

`has_at` 変数には `'@' in email` という値を設定してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_at').is_equivalent("has_at = '@' in email")
`) })
```

`has_at` を出力して、結果が `Has @ symbol: True` と表示されるようにしてください。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has @ symbol: {has_at}')",
    "print('Has @ symbol: ' + str(has_at))",
    "print('Has @ symbol:', has_at)",
    "print('Has @ symbol:', str(has_at))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

--fcc-editable-region--

--fcc-editable-region--
```
