---
id: 69e62ce96df237884da25282
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

`.lower()` メソッドは文字列内のすべての文字を小文字に変換します。これはメールアドレスを正規化するのに便利です。

```python
upperGreet = 'HELLO WORLD'
lowerGreet = upperGreet.lower()
print(lowerGreet)  # Output: hello world
```

文字列に対して `.strip()` と `.lower()` の両方をこのように呼び出せます。

```python
messy = '  HELLO WORLD  '
clean = messy.strip().lower()
print(clean)  # Output: hello world
```

変数 `messy_email` を作成し、先頭にスペース2つ、末尾にスペース2つを含む文字列 `SARAH.DAVIS@COMPANY.COM` を代入してください。

さらに、`messy_email` に対して `.strip().lower()` を呼び出し、その結果を変数 `clean_email` に代入してください。

次に両方を出力して、結果が以下のようになるようにしてください。

```md
Original email: '  SARAH.DAVIS@COMPANY.COM  '
Cleaned email: 'sarah.davis@company.com'
```

# --hints--

`messy_email` という名前の変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_email')
`) })
```

`messy_email` に <code>'  SARAH.DAVIS@COMPANY.COM  '</code> を代入してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_email').is_equivalent("messy_email = '  SARAH.DAVIS@COMPANY.COM  '")
`) })
```

`clean_email` という名前の変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('clean_email')
`) })
```

`messy_email` に対して `.strip().lower()` を呼び出し、その結果を `clean_email` 変数に代入してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_email').is_equivalent("clean_email = messy_email.strip().lower()")
`) })
```

`messy_email` を出力して、結果が `Original email: '  SARAH.DAVIS@COMPANY.COM  '` となるようにしてください。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original email: '{messy_email}'\\")",
    "print(\\"Original email: '\\" + messy_email + \\"'\\")",
    "print(\\"Original email: '\\", messy_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

`clean_email` を出力して、結果が `Cleaned email: 'sarah.davis@company.com'` となるようにしてください。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned email: '{clean_email}'\\")",
    "print(\\"Cleaned email: '\\" + clean_email + \\"'\\")",
    "print(\\"Cleaned email: '\\", clean_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
