---
id: 69e6299710319e7cb8ebdc37
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

実際のデータは、余分なスペースや不揃いな大文字・小文字で乱雑になっていることがよくあります。

変数 `messy_name` を作成し、先頭と末尾にそれぞれスペースが2つずつある文字列 `sARaH dAVis` を代入してください。その後、出力が次のようになるように出力してください。

```md
Original name: '  sARaH dAVis  '
```

この後のステップでこれをきれいにします。

# --hints--

`messy_name` という名前の変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_name')
`) })
```

`messy_name` に <code>'  sARaH dAVis  '</code> を代入してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_name').is_equivalent("messy_name = '  sARaH dAVis  '")
`) })
```

`messy_name` を出力して、出力が `Original name: '  sARaH dAVis  '` となるようにしてください。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original name: '{messy_name}'\\")",
    "print(\\"Original name: '\\" + messy_name + \\"'\\")",
    "print(\\"Original name: '\\", messy_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

--fcc-editable-region--

--fcc-editable-region--
```
