---
id: 69e75dc335b6ec9903a87d3e
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

不要な文字を削除するために `.replace()` を使うこともできます。例えば、電話番号からダッシュを取り除く場合です。

```py
phone = '456-147'
clean = phone.replace('-', '')
print(clean)  # Output: 456147
```

`phone` という変数を作成し、文字列 `123-456-7890` を代入してください。次に `-` を空文字に置き換え、その結果を `phone_clean` という変数に代入してください。

その後、両方を出力して、結果が次のようになるようにしてください。

```md
Phone with dashes: 123-456-7890
Phone without dashes: 1234567890
```

# --hints--

`phone` という名前の変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('phone')
`) })
```

`phone` に文字列 `123-456-7890` を代入してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('phone').is_equivalent("phone = '123-456-7890'")
`) })
```

`phone_clean` という名前の変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('phone_clean')
`) })
```

`phone` に対して `.replace('-', '')` を使い、その結果を `phone_clean` 変数に代入してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('phone_clean').is_equivalent("phone_clean = phone.replace('-', '')")
`) })
```

`phone` を出力して、結果が `Phone with dashes: 123-456-7890` となるようにしてください。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Phone with dashes: {phone}')",
    "print('Phone with dashes: ' + phone)",
    "print('Phone with dashes:', phone)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

`phone_clean` を出力して、結果が `Phone without dashes: 1234567890` となるようにしてください。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Phone without dashes: {phone_clean}')",
    "print('Phone without dashes: ' + phone_clean)",
    "print('Phone without dashes:', phone_clean)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

--fcc-editable-region--

--fcc-editable-region--
```
