---
id: 69e736c39dfa403496e8071e
title: ステップ11
challengeType: 20
dashedName: step-11
---

# --description--

`.upper()` メソッドはすべての文字を大文字に変換します。これは名札や見出しに便利です。

```py
name = 'sarah davis'
badge = name.upper()
print(badge)  # Output: SARAH DAVIS
```

`name_badge` という変数を作成し、`clean_name` に対して `.upper()` メソッドを呼び出した結果を代入してください。

次に `name_badge` を出力して、結果が `Name badge: SARAH DAVIS` と表示されるようにしてください。

# --hints--

`name_badge` という変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('name_badge')
`) })
```

`clean_name` に対して `.upper()` を呼び出し、その結果を `name_badge` 変数に代入してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('name_badge').is_equivalent("name_badge = clean_name.upper()")
`) })
```

`name_badge` を出力して、結果が `Name badge: SARAH DAVIS` と表示されるようにしてください。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Name badge: {name_badge}')",
    "print('Name badge: ' + name_badge)",
    "print('Name badge:', name_badge)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

--fcc-editable-region--

--fcc-editable-region--
```
