---
id: 69e73f4f5f986657647bbd63
title: ステップ 13
challengeType: 20
dashedName: step-13
---

# --description--

`.replace(old, new)` メソッドは、部分文字列 `old` のすべての出現箇所を別の文字列 `new` に置き換えます。他のメソッドと連結して使うこともできます。例えば：

```py
text = 'john.doe'
name = text.replace('.', ' ').title()
print(name)  # Output: John Doe
```

`email_user` の中の `.` を半角スペースに置き換え、その結果に対して `.title()` を呼び出してください。結果を `display_name` という変数に代入してください。

次に `display_name` を出力して、表示が `Display name: Sarah Davis` になるようにしてください。

# --hints--

`display_name` という変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('display_name')
`) })
```

`email_user` に対して `.replace('.', ' ')` と `.title()` を使い、その結果を `display_name` 変数に代入してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('display_name').is_equivalent("display_name = email_user.replace('.', ' ').title()")
`) })
```

`display_name` を出力して、表示が `Display name: Sarah Davis` になるようにしてください。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Display name: {display_name}')",
    "print('Display name: ' + display_name)",
    "print('Display name:', display_name)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

--fcc-editable-region--

--fcc-editable-region--
```
