---
id: 69e76e91402d7850510a5f18
title: ステップ 16
challengeType: 20
dashedName: step-16
---

# --description--

`.join(list)` メソッドは `.split()` の逆です。区切り文字を使って文字列のリストを1つの文字列に結合します。例えば：

```py
parts = ['apple', 'banana', 'cherry']
joined = ' - '.join(parts)
print(joined)  # Output: apple - banana - cherry
```

`address_parts` の要素を、両側にスペースがある `|` 文字で結合し、その結果を `rejoined` という変数に代入してください。

次に `rejoined` を出力して、結果が `Rejoined: 123 Main Street | Springfield | IL` と表示されるようにしてください。

# --hints--

`rejoined` という変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('rejoined')
`) })
```

`rejoined` に `' | '.join(address_parts)` の値を代入してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('rejoined').is_equivalent("rejoined = ' | '.join(address_parts)")
`) })
```

`rejoined` を出力して、結果が `Rejoined: 123 Main Street | Springfield | IL` と表示されるようにしてください。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Rejoined: {rejoined}')",
    "print('Rejoined: ' + rejoined)",
    "print('Rejoined:', rejoined)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print(f'Address parts: {address_parts}')

--fcc-editable-region--

--fcc-editable-region--
```
