---
id: 69e774b82fe0bfdb684f95b4
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

文字列がどのように始まるかを確認する必要がある場合があります。Pythonはこれを行うために`.startswith()`というメソッドを提供しています。文字列が指定した接頭辞で始まる場合は`True`を返し、そうでなければ`False`を返します。

例を示します。

```python
code = 'EMP-001'
print(code.startswith('EMP')) # True
print(code.startswith('001')) # False
```

`employee_code`という変数を作成し、文字列`EMP-2024-SD`を代入してください。次に、`employee_code`が`EMP`で始まるかどうかをチェックする`starts_with_emp`という変数を作成してください。

# --hints--

`employee_code`という名前の変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('employee_code')
`) })
```

`employee_code`に`EMP-2024-SD`を代入してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('employee_code').is_equivalent("employee_code = 'EMP-2024-SD'")
`) })
```

`starts_with_emp`という名前の変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('starts_with_emp')
`) })
```

`employee_code`に対して`.startswith('EMP')`を使い、その結果を`starts_with_emp`変数に代入してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('starts_with_emp').is_equivalent("starts_with_emp = employee_code.startswith('EMP')")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print(f'Address parts: {address_parts}')

rejoined = ' | '.join(address_parts)
print(f'Rejoined: {rejoined}')

--fcc-editable-region--

--fcc-editable-region--
```
