---
id: 69e76d1997661be7f0040e2d
title: ステップ 15
challengeType: 20
dashedName: step-15
---

# --description--

`.split(separator)` メソッドは、`separator` が現れるたびに文字列を部分文字列のリストに分割します。例えば：

```py
fruits = 'apple, banana, cherry'
parts = fruits.split(', ')
print(parts)  # Output: ['apple', 'banana', 'cherry']
```

変数 `full_address` を作成し、文字列 `123 Main Street, Springfield, IL` を代入してください。次に、カンマとそれに続く半角スペースで分割して `address_parts` を作成してください。

次に `address_parts` を出力して、結果が `Address parts: ['123 Main Street', 'Springfield', 'IL']` となるようにしてください。

# --hints--

`full_address` という名前の変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('full_address')
`) })
```

`full_address` に文字列 `123 Main Street, Springfield, IL` を代入してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('full_address').is_equivalent("full_address = '123 Main Street, Springfield, IL'")
`) })
```

`address_parts` という名前の変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('address_parts')
`) })
```

`full_address` に対して `.split(', ')` を使い、その結果を `address_parts` 変数に代入してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('address_parts').is_equivalent("address_parts = full_address.split(', ')")
`) })
```

`address_parts` を出力して、結果が `Address parts: ['123 Main Street', 'Springfield', 'IL']` となるようにしてください。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Address parts: {address_parts}')",
    "print('Address parts: ' + str(address_parts))",
    "print('Address parts:', address_parts)",
    "print('Address parts:', str(address_parts))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

--fcc-editable-region--

--fcc-editable-region--
```
