---
id: 69e73a141addc6f442924b0d
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

すでに `email` 変数で `.find()` とスライスを使いました。今度は同じ手法を `clean_email` に使います。

`clean_email` の中の `@` 文字の位置を見つけて `at_pos` にストアしてください。次にスライスを使って `@` の前のすべてを取り出し、`email_user` にストアしてください。

次に `clean_email` と `email_user` を出力して、結果が次のようになるようにしてください：

```md
Email: sarah.davis@company.com
Username: sarah.davis
```

# --hints--

`at_pos` という名前の変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_pos')
`) })
```

`clean_email` に対して `.find('@')` を呼び出し、その結果を `at_pos` 変数に代入してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_pos').is_equivalent("at_pos = clean_email.find('@')")
`) })
```

`email_user` という名前の変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email_user')
`) })
```

`clean_email` に対して `at_pos` を使ったスライスを行い、その結果を `email_user` 変数に代入してください。

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('email_user')
assert node.is_equivalent("email_user = clean_email[:at_pos]") or node.is_equivalent("email_user = clean_email[0:at_pos]")
`) })
```

`clean_email` を出力して、結果が `Email: sarah.davis@company.com` となるようにしてください。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {clean_email}')",
    "print('Email: ' + clean_email)",
    "print('Email:', clean_email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

`email_user` を出力して、結果が `Username: sarah.davis` となるようにしてください。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {email_user}')",
    "print('Username: ' + email_user)",
    "print('Username:', email_user)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

--fcc-editable-region--

--fcc-editable-region--
```
