---
id: 6a71a8df10e823e77358d896
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

`.com` が `email` に含まれているかをチェックする、`has_com` という名前の変数をもう一つ作成してください。

次に、出力が `Has .com: True` と表示されるように `has_com` を出力してください。

# --hints--

`has_com` という名前の変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_com')
`) })
```

`has_com` 変数には `'.com' in email` の値を設定してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_com').is_equivalent("has_com = '.com' in email")
`) })
```

出力が `Has .com: True` と表示されるように `has_com` を出力してください。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has .com: {has_com}')",
    "print('Has .com: ' + str(has_com))",
    "print('Has .com:', has_com)",
    "print('Has .com:', str(has_com))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')
--fcc-editable-region--

--fcc-editable-region--
```
