---
id: 6a71b3e49668c3be461257b5
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

`.title()` メソッドは各単語の最初の文字を大文字にします。例を示します。

```py
lowercase = 'hello world'
title_case = lowercase.title()
print(title_case)  # Output: Hello World
```

`clean_name` という変数を作成し、`messy_name` に対して `.title()` を呼び出した結果を代入してください。

# --hints--

`clean_name` という変数を作成してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('clean_name')
`) })
```

`messy_name` に対して `.title()` を使い、`clean_name` に値を代入してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_name').is_equivalent("clean_name = messy_name.title()")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
