---
id: 6a86b9e64da746838e67fb13
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

`.strip()` メソッドは文字列の両端から空白を削除します。

```python
messy = '  hello world  '
clean = messy.strip()
print(clean)  # Output: hello world
```

`.strip()` と `.title()` を連結して、文字列の両端から空白を削除し、各単語の最初の文字を大文字にできます。

```python
messy = '  hello world  '
clean = messy.strip().title()
print(clean)  # Output: Hello World
```

`clean_name` の値を更新して、`messy_name` に対して `.title()` の前に `.strip()` を呼び出すようにしてください。

その後、`clean_name` を出力して、結果が `Cleaned name: 'Sarah Davis'` になるようにしてください。

# --hints--

`clean_name` を更新して、`messy_name` に対して `.title()` の前に `.strip()` を呼び出すようにしてください。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_name').is_equivalent("clean_name = messy_name.strip().title()")
`) })
```

`clean_name` を出力して、結果が `Cleaned name: 'Sarah Davis'` になるようにしてください。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned name: '{clean_name}'\\")",
    "print(\\"Cleaned name: '\\" + clean_name + \\"'\\")",
    "print(\\"Cleaned name: '\\", clean_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--
clean_name = messy_name.title()
--fcc-editable-region--
```
