---
id: 6a870cc60c23f0332a2fe788
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

以前のワークショップで、文字列の中に変数を含めるために、開きの引用符の前に `f` を書き、変数を `{}` で囲むことで f-string を使えることを覚えていますか。

動作の仕組みをもう一度確認しましょう：

```python
name = 'Alice'
print(f'Hello, {name}')  # Hello, Alice
```

では、`email` を出力して、結果が `Email: alice.johnson@company.com` となるようにしてください。

# --hints--

`email` を出力して、結果が `Email: alice.johnson@company.com` となるようにしてください。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {email}')",
    "print('Email: ' + email)",
    "print('Email:', email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
--fcc-editable-region--

--fcc-editable-region--
```
