---
id: 68e51d8f163d3f8b70f9c4e5
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви попрактикуєтеся у роботі з методами `includes()` та `slice()`.

Для початку створіть змінну з ім’ям `fccSentence` і призначте їй рядок `"freeCodeCamp is a great place to learn web development."`.

# --hints--

Ви повинні оголосити змінну з ім’ям `fccSentence`.

```js
assert.exists(fccSentence);
```

Ваша змінна `fccSentence` має бути рядком.

```js
assert.isString(fccSentence);
```

Вашій змінній `fccSentence` має бути призначений рядок `"freeCodeCamp is a great place to learn web development."`.

```js
assert.strictEqual(fccSentence, "freeCodeCamp is a great place to learn web development.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
