---
id: 68e6172875008fea742a63c9
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Як ви пам’ятаєте з попередніх уроків, метод `includes()` перевіряє, чи містить рядок певний підрядок, і повертає `true` або `false`.

Ось приклад із використанням методу `includes()`:

```js
const text = "Hello World";
const hasHello = text.includes("Hello");
console.log(hasHello); // true
```

Створіть змінну з іменем `hasFreeCodeCamp`. Потім призначте їй результат використання методу `includes()` на `fccSentence`, щоб перевірити, чи містить він `"freeCodeCamp"`.

# --hints--

Ви повинні оголосити змінну з іменем `hasFreeCodeCamp`.

```js
assert.exists(hasFreeCodeCamp);
```

Ви повинні використати `fccSentence.includes("freeCodeCamp")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

Ви повинні призначити результат `fccSentence.includes("freeCodeCamp")` вашій змінній `hasFreeCodeCamp`.

```js
assert.match(__helpers.removeJSComments(code), /hasFreeCodeCamp\s*=\s*fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

Результат `fccSentence.includes("freeCodeCamp")` має бути `true`.

```js
assert.isTrue(hasFreeCodeCamp);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

--fcc-editable-region--

--fcc-editable-region--
```
