---
id: 68e6172975008fea742a63d1
title: Крок 11
challengeType: 1
dashedName: step-11
---

# --description--

Метод `slice()` може використовувати від’ємні індекси, щоб рахувати з кінця рядка. `-1` позначає останній символ.

Ось приклад:

```js
const text = "JavaScript";
const lastThree = text.slice(-3);
console.log(lastThree); // "ipt"
```

Створіть змінну з іменем `endPunctuation`. Потім призначте їй результат використання `message.slice(-1)` для вилучення останнього символу.

Потім використайте або шаблонний рядок, або конкатенацію рядків, щоб вивести повідомлення `The ending punctuation mark is a "<endPunctuation>"` у консоль. Замініть `<endPunctuation>` на фактичне значення змінної.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Ви повинні оголосити змінну з іменем `endPunctuation`.

```js
assert.exists(endPunctuation);
```

Ваша змінна `endPunctuation` має бути рядком.

```js
assert.isString(endPunctuation);
```

Ви повинні використати `message.slice(-1)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*-1\s*\)/);
```

Ви повинні призначити результат `message.slice(-1)` вашій змінній `endPunctuation`.

```js
assert.match(__helpers.removeJSComments(code), /endPunctuation\s*=\s*message\.slice\(\s*-1\s*\)/);
```

Ваша змінна `endPunctuation` має дорівнювати `"!"`.

```js
assert.strictEqual(endPunctuation, "!");
```

Ви повинні вивести повідомлення `The ending punctuation mark is a "<endPunctuation>"`, де `<endPunctuation>` має бути замінене на фактичне значення змінної.

```js
assert.deepInclude(spy.calls, ['The ending punctuation mark is a "!"']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

const greetingWord = message.slice(0, 7);
console.log(`The first word is "${greetingWord}".`);

--fcc-editable-region--

--fcc-editable-region--
```
