---
id: 68e6172975008fea742a63cb
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Тепер настав час побачити, що відбувається, коли підрядок НЕ знайдено.

Створіть змінну з іменем `hasJavaScript`. Потім призначте їй результат використання методу `includes()` для перевірки, чи містить `fccSentence` `"JavaScript"`.

Потім використайте або шаблонний рядок, або конкатенацію рядків, щоб вивести повідомлення `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.` у консоль. Замініть `<hasJavaScript>` на фактичне значення змінної.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Ви повинні оголосити змінну з іменем `hasJavaScript`.

```js
assert.exists(hasJavaScript);
```

Ви повинні використати `fccSentence.includes("JavaScript")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

Ви повинні призначити результат `fccSentence.includes("JavaScript")` вашій змінній `hasJavaScript`.

```js
assert.match(__helpers.removeJSComments(code), /hasJavaScript\s*=\s*fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

Результат `fccSentence.includes("JavaScript")` має бути `false`.

```js
assert.isFalse(hasJavaScript);
```

Ви повинні вивести повідомлення `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.`, де `<hasJavaScript>` слід замінити на фактичне значення змінної.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("JavaScript") returns false because the word "JavaScript" is not in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
