---
id: 68e6172975008fea742a63ce
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Метод `slice()` витягує розділ рядка. Він приймає два аргументи: початковий індекс (включно) та кінцевий індекс (не включно).

Ось приклад:

```js
const greeting = "Hello World";
const firstWord = greeting.slice(0, 5);
console.log(firstWord); // "Hello"
```

Індекси рядків починаються з 0. У рядку `message` символ з індексом 11 — це `f` (початок "freeCodeCamp").

Створіть змінну з ім’ям `platform`. Потім призначте їй результат використання `message.slice(11, 23)` для витягування `"freeCodeCamp"`.

# --hints--

Ви повинні оголосити змінну з ім’ям `platform`.

```js
assert.exists(platform);
```

Ваша змінна `platform` має бути рядком.

```js
assert.isString(platform);
```

Ви повинні використати `message.slice(11, 23)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*11\s*,\s*23\s*\)/);
```

Ви повинні призначити результат `message.slice(11, 23)` вашій змінній `platform`.

```js
assert.match(__helpers.removeJSComments(code), /platform\s*=\s*message\.slice\(\s*11\s*,\s*23\s*\)/);
```

Ваша змінна `platform` має дорівнювати `"freeCodeCamp"`.

```js
assert.strictEqual(platform, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

--fcc-editable-region--

--fcc-editable-region--
```
