---
id: 68e6172975008fea742a63cc
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Метод `includes()` чутливий до регістру, тобто розрізняє великі та малі літери.

Створіть змінну з іменем `hasLowercaseFCC`. Потім призначте їй результат використання методу `includes()` для перевірки, чи містить `fccSentence` `"freecodecamp"` (усі в нижньому регістрі).

Потім використайте або шаблонний рядок, або конкатенацію рядків, щоб вивести повідомлення `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.` у консоль. Замініть `<hasLowercaseFCC>` на фактичне значення змінної.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Ви повинні оголосити змінну з іменем `hasLowercaseFCC`.

```js
assert.exists(hasLowercaseFCC);
```

Ви повинні використати `fccSentence.includes("freecodecamp")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

Ви повинні призначити результат `fccSentence.includes("freecodecamp")` вашій змінній `hasLowercaseFCC`.

```js
assert.match(__helpers.removeJSComments(code), /hasLowercaseFCC\s*=\s*fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

Результат `fccSentence.includes("freecodecamp")` має бути `false`.

```js
assert.isFalse(hasLowercaseFCC);
```

Ви повинні вивести повідомлення, що містить `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.`, де `<hasLowercaseFCC>` слід замінити на фактичне значення змінної.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freecodecamp") returns false because includes is case-sensitive.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
