---
id: 68e6172975008fea742a63cd
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Тепер ви дослідите метод `slice()`, який витягує частину рядка.

Створіть змінну з ім’ям `message` і призначте їй рядок `"Welcome to freeCodeCamp!"`.

Потім використайте `console.log()`, щоб вивести `"Here are some examples of the slice() method:"` у консоль.

# --hints--

Ви повинні оголосити змінну з ім’ям `message`.

```js
assert.exists(message);
```

Ваша змінна `message` має бути рядком.

```js
assert.isString(message);
```

Ваша змінна `message` має бути `"Welcome to freeCodeCamp!"`.

```js
assert.strictEqual(message, "Welcome to freeCodeCamp!");
```

Ви повинні вивести `"Here are some examples of the slice() method:"` у консоль.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*('|"|`)Here are some examples of the slice\(\) method:\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

--fcc-editable-region--

--fcc-editable-region--
```
