---
id: 68e6172975008fea742a63ca
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Тепер використайте або шаблонний рядок, або конкатенацію рядків, щоб вивести повідомлення `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.` у консоль. Замініть `<hasFreeCodeCamp>` на фактичне значення змінної.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Ви повинні вивести повідомлення `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.`, де `<hasFreeCodeCamp>` потрібно замінити на фактичне значення змінної.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freeCodeCamp") returns true because the word "freeCodeCamp" is in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
--fcc-editable-region--

--fcc-editable-region--
```
