---
id: 68e6172975008fea742a63d0
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

Ви можете використати `slice()` для вилучення з початку рядка.

Створіть змінну з іменем `greetingWord`. Потім призначте їй результат використання `message.slice(0, 7)` для вилучення першого слова `"Welcome"`.

Потім використайте або шаблонний рядок, або конкатенацію рядків, щоб вивести повідомлення `The first word is "<greetingWord>".` у консоль. Замініть `<greetingWord>` на фактичне значення змінної.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Ви повинні оголосити змінну з іменем `greetingWord`.

```js
assert.exists(greetingWord);
```

Ваша змінна `greetingWord` має бути рядком.

```js
assert.isString(greetingWord);
```

Ви повинні використати `message.slice(0, 7)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*0\s*,\s*7\s*\)/);
```

Ви повинні призначити результат `message.slice(0, 7)` вашій змінній `greetingWord`.

```js
assert.match(__helpers.removeJSComments(code), /greetingWord\s*=\s*message\.slice\(\s*0\s*,\s*7\s*\)/);
```

Ваша змінна `greetingWord` має дорівнювати `"Welcome"`.

```js
assert.strictEqual(greetingWord, "Welcome");
```

Ви повинні вивести повідомлення `The first word is "<greetingWord>".`, де `<greetingWord>` має бути замінений на фактичне значення змінної.

```js
assert.deepInclude(spy.calls, ['The first word is "Welcome".']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

--fcc-editable-region--

--fcc-editable-region--
```
