---
id: 68ef6e42fa6e8a9a18a401fa
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Використайте `console.log()`, щоб вивести `"Here are some examples of the includes() method:"` у консоль.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Ви маєте вивести `"Here are some examples of the includes() method:"` у консоль.

```js
assert.deepInclude(spy.calls, ["Here are some examples of the includes() method:"])
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

--fcc-editable-region--

--fcc-editable-region--
```
