---
id: 683ecb010f4c6bafd428802c
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個工作坊中，你將驗證一組醫療資料，以確保它符合一組規則。

宣告一個名為 `medical_records` 的變數並指定一個空列表。接下來的步驟中，你將使用它來儲存你的醫療資料。

# --hints--

你應該宣告一個名為 `medical_records` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable("medical_records")
`) })
```

你應該為你的 `medical_records` 變數指定一個空的列表。

```js
({ test: () => runPython(`
assert medical_records == []
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

