---
id: 6842c3eee327042f1da17b89
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

創建一個 `if` 述語。對於它的條件，使用 `not` 運算子來否定 `is_sequence`。在 `if` 述語內，列印 `Invalid format: expected a list or tuple.` 並傳回 `False`。

# --hints--

你應該在你的 `validate` 函式（程式）中創建一個 `if` 述語。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_ifs()[0]
`) })
```

你的 `if` 述語應該以 `not is_sequence` 作為條件。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_ifs()[0].find_conditions()[0].is_equivalent("not is_sequence")
`) })
```

你應該在你的 `if` 述語中列印 `Invalid format: expected a list or tuple.`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_ifs()[0].find_bodies()[0].has_call("print('Invalid format: expected a list or tuple.')")
`) })
```

你應該在你的 `if` 述語結尾傳回 `False`。

```js
({ test: () => runPython(`
_if_body = _Node(_code).find_function("validate").find_ifs()[0].find_bodies()[0]
_p_statement = "print('Invalid format: expected a list or tuple.')"

if not _if_body.has_call(_p_statement):
  assert _if_body[-1].is_equivalent("return False")
else:
  assert _if_body.is_ordered(_p_statement, "return False")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]
def validate(data):
    is_sequence = isinstance(data, (list, tuple))
    
--fcc-editable-region--
    
--fcc-editable-region--
```
