---
id: 6842d5b537c64266420faf62
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

為你的字典添加一個鍵 `age`，其值為整數 `34`，以及一個鍵 `gender`，其值為字串 `Female`。不要忘記鍵值對之間的逗號。

# --hints--

你的 `medical_records` 列表應該只包含一個字典。

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

你的字典中應該有一個 `age` 鍵，其值為整數 `34`。

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

你的字典中應該有一個 `gender` 鍵，其值為字串 `Female`。

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

你的字典中應該有一個 `patient_id` 鍵，其值為字串 `P1001`。

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
--fcc-editable-region--
        'patient_id': 'P1001'
        
--fcc-editable-region--
    }    
]
```
