---
id: 6842d659a4ae7973dabab3b0
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

完成字典，為你的字典添加以下三個鍵值配對：

- 一個鍵為 `diagnosis`，值為 `Hypertension`
- 一個鍵為 `medications`，其值為列表 `['Lisinopril']`
- 一個鍵為 `last_visit_id`，值為 `V2301`。

# --hints--

你應該在你的字典中有一個 `diagnosis` 鍵，其值為字串 `Hypertension`。

```js
({ test: () => runPython(`
assert medical_records[0].get("diagnosis", None) == "Hypertension"
`) })
```

你應該在你的字典中有一個 `medications` 鍵，其值為列表 `['Lisinopril']`。

```js
({ test: () => runPython(`
assert medical_records[0].get("medications", None) == ['Lisinopril']
`) })
```

你的字典中應該有一個 `last_visit_id` 鍵，其值為字串 `V2301`。

```js
({ test: () => runPython(`
assert medical_records[0].get("last_visit_id", None) == "V2301"
`) })
```

你的字典中應該有一個 `patient_id` 鍵，其值為字串 `P1001`。

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

你的字典中應該有一個 `age` 鍵，其值為整數 `34`。

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

你的字典中應該有一個 `gender` 鍵，其值為字串 `Female`。

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
--fcc-editable-region--
        'gender': 'Female'
        
--fcc-editable-region--
    }    
]
```
