---
id: 6842ba0086201089ea268c92
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

依照你在前面步驟中使用的相同結構，`medical_records` 列表已經為你填入其他病人的資料。歡迎隨時查看。

接下來你將開始撰寫用來驗證資料集的函式（程式）。創建一個名為 `validate` 的函式（程式），並帶有單一參數 `data`。

你想確保你的資料是列表或元組。因此，在 `validate` 函式中，宣告一個名為 `is_sequence` 的變數，並指定一個對 `isinstance` 的呼叫。將 `data` 作為第一個引數，並將包含 `list` 和 `tuple` 的元組作為第二個引數傳入。

# --hints--

你應該有一個名為 `validate` 的函式（程式）。

```js
({ test: () => runPython(`
assert _Node(_code).has_function("validate")
`) })
```

你的 `validate` 函式應該只有一個參數 `data`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_args("data")
`) })
```

你應該在你的 `validate` 函式（程式）中有一個名為 `is_sequence` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_variable("is_sequence")
`) })
```

你應該將 `isinstance(data, (list, tuple))` 指定到你的 `is_sequence` 變數。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("validate").find_variable("is_sequence")
assert _cond.is_equivalent("is_sequence = isinstance(data, (list, tuple))") or _cond.is_equivalent("is_sequence = isinstance(data, (tuple, list))")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

--fcc-editable-region--

--fcc-editable-region--
```
