---
id: 6842aee5bcaf4ab8883f8191
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

如你在先前課程中所學，字典是一種資料結構，用來儲存鍵值配對，其中鍵是不可變的資料型別（例如字串），且在字典中是唯一的：

```py
person = {
    'name': 'John',
    'age': 33
}
```

`medical_records` 列表將儲存字典，每個字典表述一位病人。為你的 `medical_records` 列表添加一個字典，其鍵為 `patient_id`，值為字串 `P1001`。

# --hints--

你的 `medical_records` 列表應該只包含一個字典。

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

你的字典中應該有一個 `patient_id` 鍵，其值為字串 `P1001`。

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
--fcc-editable-region--
    
--fcc-editable-region--
]
```
