---
id: 6846e3e61271e1f164636b13
title: 第 21 步
challengeType: 20
dashedName: step-21
---

# --description--

現在你將使驗證更細緻。創建一個名為 `find_invalid_records` 的函式（程式），用來尋找字典中的無效值。為它提供以下參數：`patient_id`、`age`、`gender`、`diagnosis`、`medications`、`last_visit_id`。

在你的新函式（程式）內，建立一個名為 `constraints` 的空字典。然後，從你的新函式（程式）傳回 `constraints`。

# --hints--

你應該有一個名為 `find_invalid_records` 的函式。

```js
({ test: () => runPython(`
assert _Node(_code).has_function("find_invalid_records")
`) })
```

你的 `find_invalid_records` 函式（程式）應該依序有以下參數：`patient_id`、`age`、`gender`、`diagnosis`、`medications`、`last_visit_id`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("find_invalid_records").has_args("patient_id, age, gender, diagnosis, medications, last_visit_id")
`) })
```

你應該在你的新函式中有一個名為 `constraints` 的變數，其值為空字典。

```js
({ test: () => assert(runPython(`
_foo = _Node(_code).find_function("find_invalid_records")
_foo.has_stmt("constraints = {}") or _foo.has_stmt("constraints = dict()")
`)) })
```

你應該從你的 `find_invalid_records` 函式傳回 `constraints` 字典。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("find_invalid_records").has_return("constraints")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

--fcc-editable-region--

--fcc-editable-region--

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

        if set(dictionary.keys()) != key_set:
            print(
                f'Invalid format: {dictionary} at position {index} has missing and/or invalid keys.'
            )
            is_invalid = True

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
```
