---
id: 6846ac424529900e7327d326
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

如你在先前的課程中所學，集合是一組無序且唯一元素的群集：

```py
integers = set([3, 5, 1, 2, 1, 3, 4])
print(integers) # {1, 2, 3, 4, 5}
```

你將使用集合來確保每個字典不包含額外或拼寫錯誤的鍵。

在 `validate` 函式（程式）內，使用 `set()` 建構式來從以下每個字典應該擁有的鍵的列表中創建一個集合：`['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']`。將該集合指定給名為 `key_set` 的變數。

# --hints--

你應該在 `validate` 函式（程式）內有一個名為 `key_set` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_variable("key_set")
`) })
```

你應該將 `set(['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id'])` 指定到你的 `key_set` 變數。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_variable("key_set").is_equivalent("key_set = set(['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id'])")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }   
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
--fcc-editable-region--
    
--fcc-editable-region--

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
```
