---
id: 684fdb4386b473631d8e27ca
title: 步驟 27
challengeType: 20
dashedName: step-27
---

# --description--

正規表達式可以包含由反斜線（`\`）後接一個字元所組成的特殊序列。這些序列具有特殊意義。例如，`\d` 會匹配一個十進位數字。

在 Python 中，為字串字面值加上 `r` 或 `R` 前置式會創建一個*未經處理的字串*。未經處理的字串會將反斜線視為字元本身，而非跳脫序列，這使它們非常適合用於正規表達式。

```py
import re

book = "Fahrenheit 451"
print(re.search(r'\d', book))
# <re.Match object; span=(11, 12), match='4'>
```

在字母 `p` 之後，`patient_id` 應該有一連串的數字。因此，將你的正規表達式樣式修改為未經處理的字串 `r'p\d'`。

# --hints--

你應該將你的正規表達式樣式修改為 `r'p\d'`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("find_invalid_records").find_variable("constraints").is_equivalent(r"constraints = {'patient_id': isinstance(patient_id, str) and re.search(r'p\\d', patient_id, re.IGNORECASE)}")
`) })
```

你應該為你的正規表達式樣式使用未經處理的字串。

```js
({ test: () => runPython(`
import ast

_target = _Node(_code).find_function("find_invalid_records").find_variable("constraints").tree.value
_keys = [k.value for k in _target.keys]
_index = _keys.index('patient_id')
_val = _target.values[_index]
_pattern = _val.values[1].args[0]
assert ast.get_source_segment(_code, _pattern).lower().startswith('r')
`) })
```

# --seed--

## --seed-contents--

```py
import re

medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def find_invalid_records(
    patient_id, age, gender, diagnosis, medications, last_visit_id
):
    constraints = {
--fcc-editable-region--
        'patient_id': isinstance(patient_id, str) and re.search('p', patient_id, re.IGNORECASE)
--fcc-editable-region--
    }
    return constraints

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

        if set(dictionary.keys()) != key_set:
            print(
                f'Invalid format: {dictionary} at position {index} has missing and/or invalid keys.'
            )
            is_invalid = True

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
print(find_invalid_records(**medical_records[0]))
```
