---
id: 6718d2d59337c822ecb697ff
title: 创建一个银行账户管理程序
challengeType: 26
dashedName: build-a-bank-account-management-program
saveSubmissionToDB: true
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。 

**用户故事：**

1. 你应该使用构造函数定义一个名为 `BankAccount` 的类。 构造函数应将默认余额 `balance` 设置为 `0`，并初始化一个名为 `transactions` 的空数组，用于以对象形式存储交易记录。

2. 存储在 `transactions` 数组中的每笔交易记录必须是一个包含两个属性的对象：`type` 和 `amount`。 属性 `type` 必须为 `deposit` 或 `withdraw`，而属性 `amount` 应记录其存入或取出的金额。

3. 你应该定义一个名为 `deposit` 的方法，它将存款金额作为一个参数。 当存款金额大于 `0`时，它应该：

   - 向交易数组推送一个类型为 `deposit` 的新对象，并包含对应的存款金额。
   - 更新余额。
   - 返回 `"Successfully deposited $[amount]. New balance: $[balance]"`。

4. 如果金额小于或等于 `0`，`deposit` 方法应该返回 `"Deposit amount must be greater than zero."`。

5. 你应该定义一个名为 `withdraw` 的方法，它将存款金额作为一个参数。 这个方法应该根据提款情况更新当前余额。 当待提取的金额大于 `0` 并且小于或等于当前余额时，它应该：

   - 向交易数组推送一个类型为 `withdraw` 的新对象，并包含对应的存款金额。
   - 更新余额。
   - 返回 `"Successfully deposited $[amount]. New balance: $[balance]"`。

6. 如果要提取的金额小于或等于 `0` 或大于当前余额。 `withdraw` 方法应该返回 `"Insufficient balance or invalid amount."`。


7. 你需定义一个名为 `checkBalance` 的方法，该方法返回格式为 `"Current balance: $[balance]"` 的当前余额信息。

8. 你需定义一个名为 `listAllDeposits` 的方法，该方法​​遍历交易数组​​，并以格式 `"Deposits: amount,amount,..."` 返回所有存款记录。

9. 你需定义一个名为 `listAllWithdrawals` 的方法，该方法​​遍历交易数组​​，并以格式 `"Withdrawals: amount,amount,..."` 返回所有取款记录。

10. 你应该创建一个 `BankAccount` 的实例，名为 `myAccount`。

11. 你的 `myAccount` 银行帐户至少应该有五笔交易。

12. 你的 `myAccount` 银行帐户至少应该有两笔存款。

13. 你的 `myAccount` 银行帐户至少应该有两笔取款。

14. 你的 `myAccount` 银行账户的余额应该大于 `$100`。

# --hints--

你应该定义一个名为 `BankAccount` 的类。

```js
assert.isFunction(BankAccount);
assert.isObject(new BankAccount());
```

`BankAccount` 对象初始应将 `balance` 设置为 `0`，并创建一个空数组 `transactions` 用于存储交易记录。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.balance, 0);
assert.equal(myAccount.transactions.length, 0);
```

你应该有一个 `deposit` 方法，将存款金额作为一个参数。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.deposit);
assert.lengthOf(myAccount.deposit, 1);
```

你应该有一个 `withdraw` 方法，将提款金额作为一个参数。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.withdraw);
assert.lengthOf(myAccount.withdraw, 1);
```

你应该有一个 `checkBalance` 方法，检查当前余额。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.checkBalance);
```

你应该有一个 `listAllDeposits` 方法，列出所有存款。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllDeposits);
```

你应该有一个 `listAllWithdrawals` 方法，列出所有提款。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllWithdrawals);
```

`BankAccount.deposit(100)` 应该返回`"Successfully deposited $100. New balance: $100"`。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(100), "Successfully deposited $100. New balance: $100")

// prevent hardcoding
assert.equal(myAccount.deposit(200), "Successfully deposited $200. New balance: $300")

```

`BankAccount.deposit(-50)` 应该返回 `"Deposit amount must be greater than zero."`。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(-50), "Deposit amount must be greater than zero.")

// prevent hardcoding
assert.equal(myAccount.deposit(-30), "Deposit amount must be greater than zero.")

```

`BankAccount.deposit(0)` 应该返回 `"Deposit amount must be greater than zero."`。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(0), "Deposit amount must be greater than zero.")
```

当账户余额为 `100`, `BankAccount.withdraw(150)` 应该返回 `"Insufficient balance or invalid amount."`。

```js
const myAccount = new BankAccount();
myAccount.deposit(100);
assert.equal(myAccount.withdraw(150), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(250), "Insufficient balance or invalid amount.")

```

`BankAccount.withdraw(-50)` 应该返回 `"Insufficient balance or invalid amount."`。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(-50), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(-30), "Insufficient balance or invalid amount.")
```

`BankAccount.withdraw(0)` 应该返回 `"Insufficient balance or invalid amount."`。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(0), "Insufficient balance or invalid amount.")
```

当账户余额为 `200`, `BankAccount.withdraw(150)` 应该返回 `"Successfully withdrew $150. New balance: $50"`。

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.withdraw(150), "Successfully withdrew $150. New balance: $50")
```

当账户余额为 `200`, `BankAccount.checkBalance()` 应该返回 `"Current balance: $200"`。

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.checkBalance(), "Current balance: $200")

// prevent hardcoding
myAccount.deposit(300);
assert.equal(myAccount.checkBalance(), "Current balance: $500")
```

当你存入 `10`、`35`、`90`，`listAllDeposits` 方法应该返回 `"Deposits: 10,35,90"`。

```js
const myAccount = new BankAccount();
myAccount.deposit(10);
myAccount.deposit(35);
myAccount.deposit(90);
assert.equal(myAccount.listAllDeposits(), "Deposits: 10,35,90")
```

当你取出 `20`、`50`、`100`，`listAllWithdrawals` 方法应该返回 `"Withdrawals: 20,50,100"`。

```js
const myAccount = new BankAccount();
myAccount.deposit(500);

myAccount.withdraw(20);
myAccount.withdraw(50);
myAccount.withdraw(100);
assert.equal(myAccount.listAllWithdrawals(), "Withdrawals: 20,50,100")
```

你应该有一个名为 `myAccount` 的 `BankAccount` 实例。

```js
assert.exists(myAccount);
assert.instanceOf(myAccount, BankAccount);
```

你的 `myAccount` 银行帐户至少应该有五笔交易。

```js
assert.isAtLeast(myAccount.transactions.length, 5);
```

你的 `myAccount` 银行帐户至少应该有两笔存款。

```js
const depositCount = myAccount.transactions.filter(transaction => transaction.type === "deposit").length;
assert.isAtLeast(depositCount, 2);
```

你的 `myAccount` 银行帐户至少应该有两笔取款。

```js
const withdrawalCount = myAccount.transactions.filter(transaction => transaction.type === "withdraw").length;
assert.isAtLeast(withdrawalCount, 2);
```


你的 `myAccount` 银行账户的余额应该大于 `$100`。

```js
assert.isAbove(myAccount.balance, 100);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class BankAccount {
    constructor() {
        this.balance = 0;
        this.transactions = [];
    }

    deposit(amount) {
        if (amount > 0) {
            this.transactions.push({ type: 'deposit', amount });
            this.balance += amount;
            return `Successfully deposited $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Deposit amount must be greater than zero.';
        }
    }

    withdraw(amount) {
        if (amount > 0 && amount <= this.balance) {
            this.transactions.push({ type: 'withdraw', amount });
            this.balance -= amount;
            return `Successfully withdrew $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Insufficient balance or invalid amount.';
        }
    }

    checkBalance() {
        return `Current balance: $${this.balance}`;
    }

    listAllDeposits() {
        const deposits = this.transactions
            .filter(transaction => transaction.type === 'deposit')
            .map(transaction => transaction.amount);
        return `Deposits: ${deposits}`;
    }

    listAllWithdrawals() {
        const withdrawals = this.transactions
            .filter(transaction => transaction.type === 'withdraw')
            .map(transaction => transaction.amount);
        return `Withdrawals: ${withdrawals}`;
    }
}

const myAccount = new BankAccount();
  
console.log(myAccount.deposit(0)); 
console.log(myAccount.deposit(20)); 
console.log(myAccount.deposit(500));
console.log(myAccount.withdraw(30));
console.log(myAccount.withdraw(55));
console.log(myAccount.withdraw(100));
console.log(myAccount.checkBalance());
console.log(myAccount.listAllDeposits());
console.log(myAccount.listAllWithdrawals());
```
