---
id: 61329b210dac0b08047fd6ab
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

`meta` 要素を続けて、`viewport` 定義はブラウザーにページのレンダリング方法を伝えます。これを含めることでモバイルでの視覚的アクセシビリティが向上し、_SEO_（検索エンジン最適化）も改善されます。

ページの `viewport` と `content` を詳述する `width` 属性を持つ `initial-scale` 定義を追加してください。

# --hints--

`meta` にもう一つ `head` 要素を作成してください。

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

`meta` に `name` の `viewport` 属性を付けてください。

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

`meta` に `content` の `width=device-width, initial-scale=1` 属性を付けてください。

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
