---
id: 6133acc353338c0bba9cb553
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

最後に、`head`の中で、`title`要素はスクリーンリーダーがページの内容を理解するのに役立ちます。さらに、これは_SEO_の重要な部分です。

ページに説明的で簡潔な`title`を付けてください。

# --hints--

`title`に`head`要素を追加してください。

```js
assert.exists(document.querySelector('head > title'));
```

`title`は60文字を超えないようにしてください。

```js
assert.isAtMost(document.querySelector('head > title')?.textContent?.length, 60);
```

`title`要素はもっと説明的にしてください。_ヒント: 少なくとも15文字_

```js
assert.isAtLeast(document.querySelector('head > title')?.textContent?.length, 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
