---
id: 613e2546d0594208229ada50
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

`header` の中に、`img`、`h1`、`nav` の各要素をネストして、ページのコンテキストを提供してください。

`img` は `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg` を指し、`id` は `logo` で、`alt` のテキストは `freeCodeCamp` にしてください。

`h1` にはテキスト `HTML/CSS Quiz` を含めてください。

# --hints--

`img` 要素に `header` 要素を追加してください。

```js
assert.exists(document.querySelector('header > img'));
```

`h1` 要素に `header` 要素を追加してください。

```js
assert.exists(document.querySelector('header > h1'));
```

`nav` 要素に `header` 要素を追加してください。

```js
assert.exists(document.querySelector('header > nav'));
```

`img`、`h1`、`nav` の要素をその順番で配置してください。

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

`img` 要素に `src` 属性を `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg` として付与してください。

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

`img` 要素に `id` 属性を `logo` として付与してください。

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

`img` 要素に `alt` 属性を `freeCodeCamp` として付与してください。

```js
assert.equal(document.querySelector('img')?.alt, 'freeCodeCamp');
```

`h1` 要素にテキスト `HTML/CSS Quiz` を設定してください。

```js
assert.include(document.querySelector('h1')?.innerText.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
