---
id: 6140827cff96e906bd38fc2b
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

<a href="https://design-style-guide.freecodecamp.org/" target="_blank" rel="noopener noreferrer nofollow">freeCodeCampスタイルガイド</a>に記載されているように、ロゴは`35 / 4`のアスペクト比を維持し、テキストの周りにパディングを持つ必要があります。

まず、`background-color` を `#0a0a23` に変更してロゴが見えるようにしてください。次に、`aspect-ratio` プロパティを使って希望のアスペクト比を `35 / 4` に設定します。最後に、全体に `padding` を `0.4rem` 追加してください。

# --hints--

`#logo`に`background-color`として`#0a0a23`を設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.backgroundColor, 'rgb(10, 10, 35)');
```

`aspect-ratio`プロパティを使ってください。

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio);
```

`height`プロパティは使わないでください。

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.height);
```

`width`プロパティは変更しないでください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

`img`に`35 / 4`の`aspect-ratio`を設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio, '35 / 4');
```

`img` に `padding` の `0.4rem` を設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.padding, '0.4rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

#logo {
  width: max(10rem, 18vw);
  --fcc-editable-region--
  
  --fcc-editable-region--
}

```
