---
id: 6141fabd6f75610664e908fd
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

これはクイズなので、ユーザーが回答を提出できるフォームが必要です。フォーム内のコンテンツは`section`要素を使って意味的に区切ることができます。

`main`要素の中に、3つの`section`要素を含むフォームを作成してください。そして、そのフォームが正しいメソッドを使って`https://freecodecamp.org/practice-project/accessibility-quiz`に送信されるようにしてください。

# --hints--

`form`要素の中に`main`要素をネストしてください。

```js
assert.exists(document.querySelector('main > form'));
```

`section`要素の中に3つの`form`要素をネストしてください。

```js
assert.equal(document.querySelectorAll('main > form > section')?.length, 3);
```

`form`要素に`action`属性を付けてください。

```js
assert.notEmpty(document.querySelector('main > form')?.action);
```

`action`属性に`https://freecodecamp.org/practice-project/accessibility-quiz`の値を付けてください。

```js
assert.equal(document.querySelector('main > form')?.action, 'https://freecodecamp.org/practice-project/accessibility-quiz');
```

`form`要素に`method`属性を付けてください。

```js
assert.notEmpty(document.querySelector('main > form')?.method);
```

`form`要素に`method`属性を付け、値は`post`にしてください。

```js
assert.equal(document.querySelector('main > form')?.method?.toLowerCase(), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
--fcc-editable-region--
    <main>

    </main>
--fcc-editable-region--
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
