---
id: 614202874ca576084fca625f
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--

すべての`region`役割にはラベルが必要で、これはスクリーンリーダーユーザーがその領域の目的を理解するのに役立ちます。ラベルを追加する1つの方法は、その領域内に見出し要素を追加し、`aria-labelledby`属性で参照することです。

次の`aria-labelledby`属性を`section`要素に追加してください：

- `student-info`
- `html-questions`
- `css-questions`

それから、各`section`要素内に、対応する`h2`属性と一致する`id`を持つ`aria-labelledby`要素をネストしてください。各`h2`に適切なテキストコンテンツを与えてください。

# --hints--

最初の`section`要素には`aria-labelledby`属性として`student-info`を与えるべきです。

```js
assert.equal(document.querySelectorAll('section')?.[0]?.getAttribute('aria-labelledby'), 'student-info');
```

2番目の`section`要素には`aria-labelledby`属性として`html-questions`を与えるべきです。

```js
assert.equal(document.querySelectorAll('section')?.[1]?.getAttribute('aria-labelledby'), 'html-questions');
```

3番目の`section`要素には`aria-labelledby`属性として`css-questions`を与えるべきです。

```js
assert.equal(document.querySelectorAll('section')?.[2]?.getAttribute('aria-labelledby'), 'css-questions');
```

最初の`h2`要素内に1つの`section`要素をネストするべきです。

```js
assert.equal(document.querySelectorAll('section')?.[0]?.querySelectorAll('h2')?.length, 1);
```

2番目の`h2`要素内に1つの`section`要素をネストするべきです。

```js
assert.equal(document.querySelectorAll('section')?.[1]?.querySelectorAll('h2')?.length, 1);
```

3番目の`h2`要素内に1つの`section`要素をネストするべきです。

```js
assert.equal(document.querySelectorAll('section')?.[2]?.querySelectorAll('h2')?.length, 1);
```

最初の`h2`要素には`id`属性として`student-info`を与えるべきです。

```js
assert.equal(document.querySelectorAll('h2')?.[0]?.id, 'student-info');
```

2番目の`h2`要素には`id`属性として`html-questions`を与えるべきです。

```js
assert.equal(document.querySelectorAll('h2')?.[1]?.id, 'html-questions');
```

3番目の`h2`要素には`id`属性として`css-questions`を与えるべきです。

```js
assert.equal(document.querySelectorAll('h2')?.[2]?.id, 'css-questions');
```

最初の`h2`要素には適切なテキストコンテンツを与えるべきです。_ヒント: 私なら`Student Info`を選びます__

```js
assert.isNotEmpty(document.querySelectorAll('h2')[0]?.textContent.trim());
```

2番目の`h2`要素には適切なテキストコンテンツを与えるべきです。_ヒント: 私なら`HTML`を選びます__

```js
assert.isNotEmpty(document.querySelectorAll('h2')[1]?.textContent.trim());
```

3番目の`h2`要素には適切なテキストコンテンツを与えるべきです。_ヒント: 私なら`CSS`を選びます__

```js
assert.isNotEmpty(document.querySelectorAll('h2')[2]?.textContent.trim());
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region"></section>
        <section role="region"></section>
        <section role="region"></section>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
