---
id: 614ccc21ea91ef1736b9b578
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、基本的なボイラープレートが提供されています。

アクセシビリティの旅を始めるために、`lang` 要素に `html` 属性を付けてください。これにより、スクリーンリーダーがページの言語を識別しやすくなります。

また、スタイルシートをページにリンクすることも忘れないでください。

# --hints--

`html` 要素に `lang` 属性を付ける必要があります。_ヒント: 英語の場合は `en` の値を使えます。_

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

`link` 要素の中に `head` 要素があるはずです。

```js
assert.exists(document.querySelector("head>link"));
```

`link` 要素には `rel="stylesheet"` と `href="styles.css"` が必要です。

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>

  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
