---
id: a202eed8fc186c8434cb6d61
title: 문자열 뒤집기 함수 만들기
challengeType: 26
dashedName: build-a-string-inverter
---

# --description--

이 실습에서는 주어진 문자열의 문자를 거꾸로 뒤집는 간단한 문자열 뒤집기 함수를 만드세요.

예를 들어, `"hello"`는 `"olleh"`가 되어야 합니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `reverseString`라는 이름의 함수를 만들고 문자열을 인수로 받아야 합니다.
2. 함수는 뒤집힌 문자열을 반환해야 합니다.

# --hints--

`reverseString`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(reverseString);
```

`reverseString`는 문자열을 인수로 받아야 합니다.

```js
assert.match(reverseString.toString(), /\s*function\s+reverseString\s*\(\s*\w+\s*\)/);
```

`reverseString("hello")`는 문자열을 반환해야 합니다.

```js
assert.isString(reverseString('hello'));
```

`reverseString("hello")`는 문자열 `olleh`를 반환해야 합니다.

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

`reverseString("Howdy")`는 문자열 `ydwoH`를 반환해야 합니다.

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

`reverseString("Greetings from Earth")`는 문자열 `htraE morf sgniteerG`를 반환해야 합니다.

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
