---
id: 68e45a2b4ab9b2f48b1cc6df
title: Diseña una página de estructura de planes de precios
challengeType: 25
dashedName: design-a-pricing-plans-layout-page
demoType: onClick
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de Usuario**

1. Tu página debe tener un elemento `h1` con el texto `Pricing Plans`.
2. Tu página debe tener un elemento `div` con la clase `pricing-container` debajo del elemento `h1`.
    - El selector `.pricing-container` debe tener una propiedad `display` con el valor `flex` y una propiedad `flex-wrap` con el valor `wrap`.
3. Dentro del elemento `.pricing-container`, debes tener tres elementos `div` con la clase `pricing-card` para representar los planes de precios.
    - Uno de los elementos `.pricing-card` debe tener la clase `basic-plan` además de la clase `pricing-card`.
    - Uno de los elementos `.pricing-card` debe tener la clase `pro-plan` además de la clase `pricing-card`.
    - Uno de los elementos `.pricing-card` debe tener la clase `premium-plan` además de la clase `pricing-card`.
4. Tu elemento `.basic-plan` debe tener un elemento `h2` con el texto `Basic`.
5. Tu elemento `.basic-plan` debe tener un elemento `p` con el texto `$9/month`.
6. Tu elemento `.pro-plan` debe tener un elemento `h2` con el texto `Pro`.
7. Tu elemento `.pro-plan` debe tener un elemento `p` con el texto `$19/month`.
8. Tu elemento `.premium-plan` debe tener un elemento `h2` con el texto `Premium`.
9. Tu elemento `.premium-plan` debe tener un elemento `p` con el texto `$29/month`.
10. Cada uno de tus elementos `.pricing-card` debe:
    - Usa Flexbox para apilar su contenido en una columna y justificar el espacio entre los hijos usando `space-between`.
    - Establece la propiedad `flex` en `0 0 200px` para darle un ancho consistente y evitar que crezca o se reduzca en el diseño.
    - Establece la propiedad `border` en `2px solid black` para ver cómo las diferentes tarjetas ocupan espacio.
11. El elemento `.basic-plan` debe aparecer primero en el diseño. Debes usar la propiedad `order` para esto.
12. El elemento `.pro-plan` debe aparecer segundo en el diseño. Debes usar la propiedad `order` y establecer su propiedad `flex-grow` en `2` para que ocupe más espacio que los otros planes.
13. El elemento `.premium-plan` debe ir al final del diseño. Debes usar la propiedad `order` para esto.

**Nota:** Asegúrate de vincular tu hoja de estilos en tu HTML y aplica tu CSS.

# --hints--

Debería tener un elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Tu página debe tener un elemento `h1` con el texto `Pricing Plans`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Pricing Plans');
```

Debes tener un elemento `div` debajo del elemento `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

Tu elemento `div` debe tener la clase `pricing-container`.

```js
assert.exists(document.querySelector('h1 + div.pricing-container'));
```

Debes tener un selector `.pricing-container`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-container'));
```

Tu selector `pricing-container` debe tener la propiedad `display` con un valor de `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.display, 'flex');
```

Tu selector `.pricing-container` debe tener la propiedad `flex-wrap` con un valor de `wrap`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.flexWrap, 'wrap');
```

Debes tener tres elementos `div` con la clase `pricing-card` dentro del elemento `.pricing-container`.

```js
const divEls = document.querySelectorAll('.pricing-container .pricing-card');
assert.lengthOf(divEls, 3);
```

Uno de tus elementos `.pricing-card` debe tener la clase `basic-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.basic-plan'));
```

Uno de tus elementos `.pricing-card` debe tener la clase `premium-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.premium-plan'));
```

Uno de tus elementos `.pricing-card` debe tener la clase `pro-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.pro-plan'));
```

Debes tener un selector `.pricing-card`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-card'));
```

Tu selector `.pricing-card` debe tener la propiedad `display` con un valor de `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.display, 'flex');
```

Tu selector `.pricing-card` debe tener la propiedad `flex-direction` con un valor de `column`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flexDirection, 'column');
```

Tu selector `.pricing-card` debe tener la propiedad `justify-content` con un valor de `space-between`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.justifyContent, 'space-between');
```

Tu selector `.pricing-card` debe tener la propiedad `flex` con un valor de `0 0 200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flex, '0 0 200px');
```

Tu selector `.pricing-card` debe tener la propiedad `border` con un valor de `2px solid black`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.border, '2px solid black');
```

Tu selector `.basic-plan` debe tener la propiedad `order` con un valor de `0`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.basic-plan')?.order, '0');
```

Tu elemento `.basic-plan` debe tener un elemento `h2` con el texto `Basic`.

```js
const h2 = document.querySelector('.basic-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Basic');
```

Tu elemento `.basic-plan` debe tener un elemento `p` con el texto `$9/month`.

```js
const p = document.querySelector('.basic-plan > p');
assert.strictEqual(p?.textContent.trim(), '$9/month');
```

Tu elemento `.pro-plan` debe tener un elemento `h2` con el texto `Pro`.

```js
const h2 = document.querySelector('.pro-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Pro');
```

Tu elemento `.pro-plan` debe tener un elemento `p` con el texto `$19/month`.

```js
const p = document.querySelector('.pro-plan > p');
assert.strictEqual(p?.textContent.trim(), '$19/month');
```

Tu elemento `.pro-plan` debe tener la propiedad `order` con un valor de `1`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.order, '1');
```

Tu elemento `.pro-plan` debe tener la propiedad `flex-grow` con un valor de `2`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.flexGrow, '2');
```

Tu elemento `.premium-plan` debe tener un elemento `h2` con el texto `Premium`.

```js
const h2 = document.querySelector('.premium-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Premium');
```

Tu elemento `.premium-plan` debe tener un elemento `p` con el texto `$29/month`.

```js
const p = document.querySelector('.premium-plan > p');
assert.strictEqual(p?.textContent.trim(), '$29/month');
```

Tu elemento `premium-plan` debe ser el último elemento en el diseño y tener la propiedad `order` con el valor `2`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.premium-plan')?.order, '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Pricing Plans Layout Page</title>
  </head>
  <body></body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Pricing Plans Layout</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Pricing Plans</h1>
    <div class="pricing-container">
      <div class="pricing-card pro-plan">
        <h2>Pro</h2>
        <p class="price">$19/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card basic-plan">
        <h2>Basic</h2>
        <p class="price">$9/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="cross">✗</span> Advanced analytics</li>
          <li><span class="cross">✗</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card premium-plan">
        <h2>Premium</h2>
        <p class="price">$29/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="check">✓</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  align-items: center;
  justify-content: center;
  flex-direction: column;
  min-height: 100vh;
}

h1 {
  margin-bottom: 20px;
}

h2,
p {
  margin: 0;
}

.pricing-container {
  display: flex;
  flex-wrap: wrap;
  align-content: center;
  border: 2px solid #000;
  gap: 20px;
  width: 80%;
  max-width: 900px;
  height: 70vh;
  padding: 0 20px;
}

.pricing-card {
  background: rgb(247, 247, 247);
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
  padding: 20px;
  flex: 0 0 200px;
  text-align: center;
  transition: transform 0.3s ease;
  display: flex;
  flex-direction: column;
  justify-content: space-between;
  border: 2px solid black;
}

.pricing-card:hover {
  transform: scale(1.05);
}

.price {
  font-size: 1.5rem;
  font-weight: bold;
  color: #3e3e3e;
  margin: 10px;
  border-bottom: 1px solid #433d28;
}

.features {
  list-style-type: none;
  padding: 0;
  text-align: left;
}

.features li {
  padding: 5px 0;
  font-size: 0.95rem;
}

.check {
  color: #219c54;
  font-weight: bold;
  margin-right: 8px;
}

.cross {
  color: crimson;
  font-weight: bold;
  margin-right: 8px;
}

a {
  text-decoration: none;
  display: inline-block;
  margin-top: 15px;
  background: #f1be32;
  color: #000;
  font-weight: bold;
  border: 3px solid #eebe3a;
  padding: 5px;
  transition: background 0.3s;
}

a:hover {
  background: #f7c436;
}

.basic-plan {
  order: 0;
}
.pro-plan {
  order: 1;
  flex-grow: 2;
  background: #f0f8ff;
}
.premium-plan {
  order: 2;
}

@media screen and (max-width: 1000px) {
  .pricing-container {
    height: auto;
    padding-bottom: 20px;
  }

  h1 {
    display: block;
    font-size: 1.5rem;
    text-align: center;
  }

  .pricing-card {
    flex: 0 0 175px;
  }
}

@media screen and (max-width: 650px) {
  .pricing-container {
    flex-direction: column;
    height: auto;
  }

  .pricing-card {
    flex: 1 1 auto;
    width: 100%;
  }
}
```

