---
id: 587d8255367417b2b2512c74
title: إنشاء صنف طابور الأولوية
challengeType: 1
forumTopicId: 301630
dashedName: create-a-priority-queue-class
---

# --description--

في هذا التحدٍّ ستنشئ طابور أولوية. طابور الأولوية هو نوع خاص من الطوابير حيث قد تحتوي العناصر على معلومات إضافية تحدد أولويتها. يمكن تمثيل ذلك ببساطة باستخدام عدد صحيح. ستتجاوز أولوية العنصر ترتيب الإدخال في تحديد متتالية إخراج العناصر من الطابور. إذا أُضيف عنصر ذو أولوية أعلى بعد عناصر ذات أولوية أقل، فسيتم إخراج العنصر ذي الأولوية الأعلى قبل جميع الآخرين.

على سبيل المثال، لنفترض أن لدينا طابور أولوية يحتوي على ثلاثة عناصر:

```js
[['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

القيمة الثانية هنا (عدد صحيح) تمثل أولوية العنصر. إذا أضفنا `['human', 1]` بأولوية `1` (مع افتراض أن الأولويات الأقل لها الأسبقية) فسيكون هذا العنصر أول من يُخرج من الطابور. ستبدو المجموعة كما يلي:

```js
[['human', 1], ['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

بدأنا بكتابة `PriorityQueue` في محرر الشفرة. ستحتاج إلى إضافة طريقة `enqueue` لإضافة عناصر مع أولوية، وطريقة `dequeue` لإزالة العناصر وإرجاعها، وطريقة `size` لإرجاع عدد العناصر في الطابور، وطريقة `front` لإرجاع العنصر في مقدمة الطابور، وأخيرًا طريقة `isEmpty` التي سترجع `true` إذا كان الطابور فارغًا أو `false` إذا لم يكن كذلك.

يجب أن تقبل طريقة `enqueue` العناصر بالتنسيق الموضح أعلاه (`['human', 1]`) حيث يمثل `1` الأولوية. يجب أن تُرجع طريقتا `dequeue` و`front` اسم العنصر فقط، وليس أولويته.

# --hints--

يجب أن يحتوي صنف `PriorityQueue` على طريقة `enqueue`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.enqueue === 'function';
  })()
);
```

يجب أن يحتوي صنف `PriorityQueue` على طريقة `dequeue`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.dequeue === 'function';
  })()
);
```

يجب أن يحتوي صنف `PriorityQueue` على طريقة `size`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.size === 'function';
  })()
);
```

يجب أن يحتوي صنف `PriorityQueue` على طريقة `front`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.front === 'function';
  })()
);
```

يجب أن يحتوي صنف `PriorityQueue` على طريقة `isEmpty`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.isEmpty === 'function';
  })()
);
```

يجب أن يتتبع صنف `PriorityQueue` عدد العناصر الحالي بشكل صحيح باستخدام طريقة `size` أثناء إضافة العناصر وإزالتها.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    test.enqueue(['Jon Snow', 1]);
    var size1 = test.size();
    test.dequeue();
    var size2 = test.size();
    test.enqueue(['A', 3]);
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    return size1 === 2 && size2 === 1 && test.size() === 4;
  })()
);
```

يجب أن تُرجع طريقة `front` العنصر الصحيح في مقدمة الطابور أثناء إضافة العناصر وإزالتها.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    var front1 = test.front();
    test.enqueue(['Jon Snow', 1]);
    var front2 = test.front();
    test.dequeue();
    test.enqueue(['A', 3]);
    var front3 = test.front();
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    test.dequeue();
    var front4 = test.front();
    return (
      front1 === 'David Brown' &&
      front2 === 'Jon Snow' &&
      front3 === 'David Brown' &&
      front4 === 'A'
    );
  })()
);
```

يجب أن تُرجع طريقة `isEmpty` القيمة `true` عندما يكون الطابور فارغًا.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 1]);
    test.enqueue(['B', 1]);
    test.dequeue();
    var first = test.isEmpty();
    test.dequeue();
    return !first && test.isEmpty();
  })()
);
```

يجب أن يُرجع طابور الأولوية العناصر ذات الأولوية الأعلى قبل العناصر ذات الأولوية الأقل، وأن يُرجع العناصر بترتيب دخولها (أول داخل، أول خارج) خلاف ذلك.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 5]);
    test.enqueue(['B', 5]);
    test.enqueue(['C', 5]);
    test.enqueue(['D', 3]);
    test.enqueue(['E', 1]);
    test.enqueue(['F', 7]);
    var result = [];
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    return result.join('') === 'EDABCF';
  })()
);
```

# --seed--

## --seed-contents--

```js
function PriorityQueue () {
  this.collection = [];
  this.printCollection = function() {
    console.log(this.collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function PriorityQueue() {
  this.collection = [];
  this.printCollection = function () {
    console.log(this.collection);
  };
  // Only change code below this line
  this.enqueue = function (newitem) {
    if (this.isEmpty()) {
      return this.collection.push(newitem);
    }

    this.collection = this.collection.reverse();
    var found_index = this.collection.findIndex(function (item) {
      return newitem[1] >= item[1];
    });
    if (found_index === -1) {
      this.collection.push(newitem);
    } else {
      this.collection.splice(found_index, 0, newitem);
    }
    this.collection = this.collection.reverse();
  };
  this.dequeue = function () {
    if (!this.isEmpty()) {
      return this.collection.shift()[0];
    } else {
      return "The queue is empty.";
    }
  };
  this.size = function () {
    return this.collection.length;
  };
  this.front = function () {
    return this.collection[0][0];
  };
  this.isEmpty = function () {
    return this.size() > 0 ? false : true;
  };
  // Only change code above this line
}
```
