---
id: 587d8256367417b2b2512c77
title: قائمة الجوار
challengeType: 1
forumTopicId: 301620
dashedName: adjacency-list
---

# --description--

يمكن تمثيل الرسوم البيانية بطرق مختلفة. هنا نصف طريقة واحدة تُسمى <dfn>قائمة الجوار</dfn>. قائمة الجوار هي في الأساس قائمة نقطية حيث الجانب الأيسر هو العقدة والجانب الأيمن يسرد كل العقد الأخرى المرتبطة بها. أدناه تمثيل لقائمة الجوار.

<blockquote>العقدة1: العقدة2، العقدة3<br>العقدة2: العقدة1<br>العقدة3: العقدة1</blockquote>

الرسم البياني أعلاه غير موجَّه لأن `Node1` متصل بـ `Node2` و`Node3`، وهذه المعلومات متوافقة مع الاتصالات التي تظهرها `Node2` و`Node3`. قائمة الجوار للرسم البياني الموجَّه تعني أن كل صف في القائمة يظهر الاتجاه. إذا كان الرسم أعلاه موجَّهًا، فإن `Node2: Node1` تعني أن الحافة الموجَّهة تشير من `Node2` نحو `Node1`. يمكننا تمثيل الرسم البياني غير الموجَّه أعلاه كقائمة جوار بوضعه داخل كائن JavaScript.

```js
var undirectedG = {
  Node1: ["Node2", "Node3"],
  Node2: ["Node1"],
  Node3: ["Node1"]
};
```

يمكن أيضًا تمثيله ببساطة كمصفوفة حيث تحمل العقد أرقامًا بدلاً من مسميات نصية.

```js
var undirectedGArr = [
  [1, 2], // Node1
  [0],    // Node2
  [0]     // Node3
];
```

# --instructions--

أنشئ شبكة اجتماعية كرسم بياني غير موجَّه يحتوي على 4 عقد/أشخاص بأسماء `James` و`Jill` و`Jenny` و`Jeff`. هناك حواف/علاقات بين James وJeff، وJill وJenny، وJeff وJenny.

# --hints--

يجب أن يحتوي `undirectedAdjList` على أربع عقد فقط.

```js
assert(Object.keys(undirectedAdjList).length === 4);
```

يجب أن تكون هناك حافة بين `Jeff` و`James`.

```js
assert(
  undirectedAdjList.James.indexOf('Jeff') !== -1 &&
    undirectedAdjList.Jeff.indexOf('James') !== -1
);
```

يجب أن تكون هناك حافة بين `Jill` و`Jenny`.

```js
assert(
  undirectedAdjList.Jill.indexOf('Jenny') !== -1 &&
    undirectedAdjList.Jenny.indexOf('Jill') !== -1
);
```

يجب أن تكون هناك حافة بين `Jeff` و`Jenny`.

```js
assert(
  undirectedAdjList.Jeff.indexOf('Jenny') !== -1 &&
    undirectedAdjList.Jenny.indexOf('Jeff') !== -1
);
```

# --seed--

## --seed-contents--

```js
var undirectedAdjList = {};
```

# --solutions--

```js
var undirectedAdjList = {
  James: ['Jeff'],
  Jill: ['Jenny'],
  Jenny: ['Jill', 'Jeff'],
  Jeff: ['James', 'Jenny']
};
```
