---
id: 587d8258367417b2b2512c81
title: حذف عقدة لها طفل واحد في شجرة البحث الثنائي
challengeType: 1
forumTopicId: 301638
dashedName: delete-a-node-with-one-child-in-a-binary-search-tree
---

# --description--

الآن بعد أن أصبح بإمكاننا حذف العقد الورقية، لننتقل إلى الحالة الثانية: حذف عقدة لها طفل واحد. في هذه الحالة، لنفترض أن لدينا شجرة تحتوي على العقد التالية 1 — 2 — 3 حيث 1 هي الجذر. لحذف 2، نحتاج ببساطة إلى جعل المرجع الأيمن في 1 يشير إلى 3. بشكل عام، لحذف عقدة لها طفل واحد فقط، نجعل عقدة الأصل تشير إلى العقدة التالية في الشجرة.

# --instructions--

لقد وفرنا بعض الشفرة في طريقة `remove` التي تنجز المهام من التحدي السابق. نحدد الهدف للحذف وعقدة الأصل له، ونحدد عدد الأطفال الذين تملكهم عقدة الهدف. لنضيف الحالة التالية هنا للعقد التي لها طفل واحد فقط. هنا، يجب أن نحدد إذا كان الطفل الوحيد هو فرع أيسر أو أيمن في الشجرة ثم نضبط المرجع الصحيح في عقدة الأصل ليشير إلى هذه العقدة. بالإضافة إلى ذلك، دعنا نأخذ في الاعتبار الحالة التي يكون فيها الهدف هو عقدة الجذر (وهذا يعني أن عقدة الأصل ستكون `null`). لا تتردد في استبدال كل شفرة البداية بشفرتك الخاصة طالما اجتازت الاختبارات.

# --hints--

يجب أن يوجد هيكل البيانات `BinarySearchTree`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

يجب أن تحتوي شجرة البحث الثنائي على طريقة تسمى `remove`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

محاولة حذف عنصر غير موجود يجب أن تُرجع `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    return test.remove(100) == null;
  })()
);
```

إذا لم يكن لعقدة الجذر أطفال، يجب أن يؤدي حذفها إلى تعيين الجذر إلى `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(500);
    test.remove(500);
    return test.inorder() == null;
  })()
);
```

يجب أن تزيل طريقة `remove` العقد الورقية من الشجرة.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(5);
    test.add(3);
    test.add(7);
    test.add(6);
    test.add(10);
    test.add(12);
    test.remove(3);
    test.remove(12);
    test.remove(10);
    return test.inorder().join('') == '567';
  })()
);
```

يجب أن تزيل طريقة `remove` العقد التي لها طفل واحد.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(1);
    test.add(4);
    test.add(3);
    test.add(2);
    test.add(6);
    test.add(8);
    test.remove(6);
    test.remove(3);
    return test.inorder().join('') == '1248';
  })()
);
```

حذف الجذر في شجرة تحتوي على عقدتين يجب أن يجعل العقدة الثانية هي الجذر.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(15);
    test.add(27);
    test.remove(15);
    return test.inorder().join('') == '27';
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    },
    inorder: function() {
      if (this.root == null) {
        return null;
      } else {
        var result = new Array();
        function traverseInOrder(node) {
          if (node.left != null) {
            traverseInOrder(node.left);
          }
          result.push(node.value);
          if (node.right != null) {
            traverseInOrder(node.right);
          }
        }
        traverseInOrder(this.root);
        return result;
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}

function BinarySearchTree() {
  this.root = null;
  this.remove = function(value) {
    if (this.root === null) {
      return null;
    }
    var target;
    var parent = null;
    // Find the target value and its parent
    (function findValue(node = this.root) {
      if (value == node.value) {
        target = node;
      } else if (value < node.value && node.left !== null) {
        parent = node;
        return findValue(node.left);
      } else if (value < node.value && node.left === null) {
        return null;
      } else if (value > node.value && node.right !== null) {
        parent = node;
        return findValue(node.right);
      } else {
        return null;
      }
    }.bind(this)());
    if (target === null) {
      return null;
    }
    // Count the children of the target to delete
    var children =
      (target.left !== null ? 1 : 0) + (target.right !== null ? 1 : 0);
    // Case 1: Target has no children
    if (children === 0) {
      if (target == this.root) {
        this.root = null;
      } else {
        if (parent.left == target) {
          parent.left = null;
        } else {
          parent.right = null;
        }
      }
    }
    // Case 2: Target has one child
    // Only change code below this line
  };
}
```

# --solutions--

```js
// solution required
```
