---
id: 587d8253367417b2b2512c6a
title: مصفوفات ذات نوع محدد
challengeType: 1
forumTopicId: 301716
dashedName: typed-arrays
---

# --description--

المصفوفات هي كائنات في JavaScript يمكنها احتواء العديد من العناصر المختلفة.

```js
var complexArr = [1, 5, "2", "Word", {"name": "James"}];
```

ما يحدث في الخلفية هو أن المتصفح يخصص تلقائيًا مقدار الذاكرة المناسب لتلك المصفوفة. كما يتغير هذا المقدار حسب الحاجة إذا أضفت أو أزلت بيانات.

مع ذلك، في عالم الأداء العالي وأنواع العناصر المختلفة، تحتاج أحيانًا إلى تحديد مقدار الذاكرة المخصصة للمصفوفة بدقة أكبر.

<dfn>المصفوفات ذات النوع المحدد</dfn> هي الحل لهذه المشكلة. يمكنك الآن تحديد مقدار الذاكرة التي تريد تخصيصها للمصفوفة. فيما يلي نظرة عامة أساسية على أنواع المصفوفات المختلفة المتاحة وحجم كل عنصر منها بالبايت.

<table><tbody><tr><th>النوع</th><th>حجم كل عنصر بالبايت</th></tr><tr><td><code>Int8Array</code></td><td>1</td></tr><tr><td><code>Uint8Array</code></td><td>1</td></tr><tr><td><code>Uint8ClampedArray</code></td><td>1</td></tr><tr><td><code>Int16Array</code></td><td>2</td></tr><tr><td><code>Uint16Array</code></td><td>2</td></tr><tr><td><code>Int32Array</code></td><td>4</td></tr><tr><td><code>Uint32Array</code></td><td>4</td></tr><tr><td><code>Float32Array</code></td><td>4</td></tr><tr><td><code>Float64Array</code></td><td>8</td></tr></tbody></table>

هناك طريقتان لإنشاء هذا النوع من المصفوفات. الطريقة الأولى هي إنشاؤها مباشرة. فيما يلي كيفية إنشاء `Int16Array` بطول 3.

```js
var i8 = new Int16Array(3);
console.log(i8);
// Returns [0, 0, 0]
```

يمكنك أيضًا إنشاء <dfn>مخزن مؤقت</dfn> لتعيين مقدار البيانات (بالبايت) التي تريد أن تشغلها المصفوفة. **ملاحظة**  
لإنشاء مصفوفات ذات نوع محدد باستخدام المخازن المؤقتة، يجب أن يكون عدد البايتات المعينة مضاعفًا للأحجام المذكورة أعلاه.

```js
// Create same Int16Array array differently
var byteSize = 6; // Needs to be multiple of 2
var buffer = new ArrayBuffer(byteSize);
var i8View = new Int16Array(buffer);
buffer.byteLength; // Returns 6
i8View.byteLength; // Returns 6
console.log(i8View); // Returns [0, 0, 0]
```

<dfn>المخازن المؤقتة</dfn> هي كائنات عامة الغرض تحمل البيانات فقط. لا يمكنك الوصول إليها بشكل مباشر. للوصول إليها، تحتاج أولًا إلى إنشاء <dfn>عرض</dfn>.

```js
i8View[0] = 42;
console.log(i8View); // Returns [42, 0, 0]
```

**ملاحظة**  
المصفوفات ذات النوع المحدد لا تحتوي على بعض الطرق التي تمتلكها المصفوفات التقليدية مثل `.pop()` أو `.push()`. كما تفشل في اختبار `Array.isArray()` الذي يتحقق مما إذا كان شيء ما مصفوفة. رغم بساطتها، يمكن أن يكون هذا ميزة لمحركات JavaScript الأقل تطورًا لتنفيذها.

# --instructions--

ابدأ بإنشاء `buffer` حجمه 64 بايتًا. ثم أنشئ مصفوفة ذات نوع محدد `Int32Array` مع عرض لها يسمى `i32View`.

# --hints--

يجب أن يكون حجم `buffer` الخاص بك 64 بايتًا.

```js
assert(buffer.byteLength === 64);
```

يجب أن يكون حجم عرض `i32View` لمخزنك المؤقت 64 بايتًا.

```js
assert(i32View.byteLength === 64);
```

يجب أن يكون طول عرض `i32View` لمخزنك المؤقت 16 عنصرًا.

```js
assert(i32View.length === 16);
```

# --seed--

## --seed-contents--

```js
var buffer;
var i32View;
```

# --solutions--

```js
var buffer = new ArrayBuffer(64);
var i32View = new Int32Array(buffer);
```
