---
id: 587d8250367417b2b2512c60
title: إنشاء صنف قائمة انتظار
challengeType: 1
forumTopicId: 301631
dashedName: create-a-queue-class
---

# --description--

مثل المكدسات، قوائم الانتظار هي مجموعة من العناصر. لكن على عكس المكدسات، تتبع قوائم الانتظار مبدأ FIFO (الدخول أولاً، الخروج أولاً). تُضاف العناصر إلى نهاية قائمة الانتظار، ويسمح فقط بإزالة العنصر الموجود في مقدمة القائمة.

يمكننا استخدام مصفوفة لتمثيل قائمة انتظار، لكن مثل المكدسات، نريد تقليل مقدار التحكم الذي نملكه على قوائم الانتظار.

الطريقتان الرئيسيتان في صنف الطابور هما طريقة enqueue وطريقة dequeue. طريقة enqueue تضيف عنصرًا إلى نهاية الطابور، وطريقة dequeue تزيل وتعيد العنصر الموجود في مقدمة الطابور. هناك طرق أخرى مفيدة مثل front وsize وisEmpty.

# --instructions--

اكتب طريقة `enqueue` تدفع عنصرًا إلى نهاية قائمة الانتظار، وطريقة `dequeue` تزيل وتُرجع العنصر الأمامي، وطريقة `front` تتيح لنا رؤية العنصر الأمامي، وطريقة `size` تعرض الطول، وطريقة `isEmpty` للتحقق مما إذا كانت قائمة الانتظار فارغة.

# --hints--

يجب أن يحتوي صنف `Queue` على طريقة `enqueue`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.enqueue === 'function';
  })()
);
```

يجب أن يحتوي صنف `Queue` على طريقة `dequeue`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.dequeue === 'function';
  })()
);
```

يجب أن يحتوي صنف `Queue` على طريقة `front`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.front === 'function';
  })()
);
```

يجب أن يحتوي صنف `Queue` على طريقة `size`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.size === 'function';
  })()
);
```

يجب أن يحتوي صنف `Queue` على طريقة `isEmpty`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.isEmpty === 'function';
  })()
);
```

يجب أن تزيل طريقة `dequeue` العنصر الأمامي من قائمة الانتظار وتُرجع قيمته.

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.dequeue() === 'Smith';
  })()
);
```

يجب أن تُرجع طريقة `front` قيمة العنصر الأمامي في قائمة الانتظار.

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.front() === 'Smith';
  })()
);
```

يجب أن تُرجع طريقة `size` طول قائمة الانتظار.

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return test.size() === 1;
  })()
);
```

يجب أن تُرجع طريقة `isEmpty` القيمة `false` إذا كانت هناك عناصر في قائمة الانتظار.

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return !test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Queue() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function Queue () { 
    var collection = [];
    this.print = function() {
        console.log(collection);
    };
    // Only change code below this line
    this.enqueue = function(item) {
        collection.push(item);
    }

    this.dequeue = function() {
        return collection.shift();
    }

    this.front = function() {
        return collection[0];
    }

    this.size = function(){
        return collection.length;
    }

    this.isEmpty = function() {
        return collection.length === 0 ? true : false;
    }
    // Only change code above this line
}
```
