---
id: 587d8255367417b2b2512c75
title: إنشاء قائمة انتظار دائرية
challengeType: 1
forumTopicId: 301625
dashedName: create-a-circular-queue
---

# --description--

في هذا التحدٍّ ستنشئ قائمة انتظار دائرية. القائمة الدائرية هي قائمة انتظار تكتب في نهاية مجموعة ثم تبدأ بالكتابة فوق نفسها في بداية المجموعة. هذا النوع من هياكل البيانات مفيد في مواقف معينة. على سبيل المثال، يمكن استخدام القائمة الدائرية لبث الوسائط. بمجرد أن تمتلئ القائمة، ستُكتب بيانات الوسائط الجديدة فوق البيانات القديمة.

طريقة جيدة لتوضيح هذا المفهوم هي باستخدام مصفوفة طولها `5`:

```js
[null, null, null, null, null]
 ^Read @ 0
 ^Write @ 0
```

هنا يكون مؤشر القراءة والكتابة كلاهما في الموضع `0`. الآن تحصل القائمة على 3 سجلات جديدة `a` و `b` و `c`. تبدو قائمتنا الآن كما يلي:

```js
[a, b, c, null, null]
 ^Read @ 0
          ^Write @ 3
```

عندما يقرأ مؤشر القراءة، يمكنه إزالة القيم أو الاحتفاظ بها:

```js
[null, null, null, null, null]
                   ^Read @ 3
                   ^Write @ 3
```

الآن نكتب القيم `d` و `e` و `f` في القائمة. بمجرد أن يصل مؤشر الكتابة إلى نهاية المصفوفة، يعود إلى البداية:

```js
[f, null, null, d, e]
                ^Read @ 3
    ^Write @ 1
```

تتطلب هذه الطريقة كمية ثابتة من الذاكرة لكنها تسمح بمعالجة ملفات بحجم أكبر بكثير.

# --instructions--

في هذا التحدٍّ سننفذ قائمة انتظار دائرية. يجب أن توفر القائمة الدائرية طريقتي `enqueue` و `dequeue` اللتين تسمحان بالقراءة والكتابة في القائمة. يجب أن يقبل الصنف نفسه معلمة صحيحة يمكنك استخدامها لتحديد حجم القائمة عند إنشائها. لقد كتبنا لك الإصدار الابتدائي من هذا الصنف في محرر الشفرة.

عند إضافة عناصر إلى الطابور، يجب أن يتقدم مؤشر الكتابة للأمام ويعود إلى البداية عند وصوله إلى نهاية الطابور. يجب أن تُرجع طريقة `enqueue` العنصر الذي أضفته إذا نجحت العملية؛ أما إذا لم تنجح، فستُرجع `null`.

بنفس الطريقة، يجب أن يتحرك مؤشر القراءة للأمام أثناء إخراج العناصر من الطابور. عند إخراج عنصر، يجب إرجاع ذلك العنصر. إذا لم تتمكن من إخراج عنصر، يجب أن تُرجع `null`.

لا يجب السماح لمؤشر الكتابة بالتقدم متجاوزًا مؤشر القراءة (فالصنف لن يسمح لك بالكتابة فوق بيانات لم تقرأها بعد) ولا يجب أن يتقدم مؤشر القراءة متجاوزًا البيانات التي كتبتها.

# --hints--

يجب أن تضيف طريقة `enqueue` العناصر إلى القائمة الدائرية.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    var print = test.print();
    return print[0] === 17 && print[1] === 32 && print[2] === 591;
  })()
);
```

يجب ألا تضيف عناصر متجاوزًا مؤشر القراءة.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    test.enqueue(13);
    test.enqueue(25);
    test.enqueue(59);
    var print = test.print();
    return print[0] === 17 && print[1] === 32 && print[2] === 591;
  })()
);
```

يجب أن تزيل طريقة `dequeue` العناصر من القائمة.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    return (
      test.dequeue() === 17 && test.dequeue() === 32 && test.dequeue() === 591
    );
  })()
);
```

بعد إزالة عنصر، يجب إعادة تعيين موضعه في القائمة إلى `null`.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(672);
    test.dequeue();
    test.dequeue();
    var print = test.print();
    return print[0] === null && print[1] === null && print[2] === 672;
  })()
);
```

محاولة إزالة عنصر متجاوزًا مؤشر الكتابة يجب أن تُرجع `null` ولا تُقدّم مؤشر الكتابة.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    return (
      test.dequeue() === 17 &&
      test.dequeue() === 32 &&
      test.dequeue() === 591 &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.enqueue(100) === 100 &&
      test.dequeue() === 100
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
class CircularQueue {
  constructor(size) {

    this.queue = [];
    this.read = 0;
    this.write = 0;
    this.max = size - 1;

    while (size > 0) {
      this.queue.push(null);
      size--;
    }
  }

  print() {
    return this.queue;
  }

  enqueue(item) {
    // Only change code below this line

    // Only change code above this line
  }

  dequeue() {
    // Only change code below this line

    // Only change code above this line
  }
}
```

# --solutions--

```js
class CircularQueue {
  constructor(size) {
    this.queue = [];
    this.read = 0;
    this.write = 0;
    this.max = size - 1;

    while (size > 0) {
      this.queue.push(null);
      size--;
    }
  }

  print() {
    return this.queue;
  }

  enqueue(item) {
    // Only change code below this line
    console.log(this.write, this.max);
    if (this.queue[this.write] === null) {
      this.queue[this.write++] = item;

      if (this.write > this.max) {
        this.write = 0;
      }
      return item;
    }
    return null;
    // Only change code above this line
  }

  dequeue() {
    // Only change code below this line
    if (this.queue[this.read] !== null) {
      let item = this.queue[this.read];
      this.queue[this.read++] = null;
      if (this.read > this.max) {
        this.read = 0;
      }
      return item;
    }
    return null;
    // Only change code above this line
  }
}
```
