---
id: 587d825b367417b2b2512c8e
title: إنشاء جدول تجزئة
challengeType: 1
forumTopicId: 301627
dashedName: create-a-hash-table
---

# --description--

في هذا التحدٍّ سنتعلم عن جداول التجزئة. يُستخدم جدول التجزئة لتنفيذ المصفوفات الترابطية، أو خرائط أزواج المفتاح-القيمة، مثل الكائنات وMaps التي درسناها للتو. يمكن تنفيذ كائن JavaScript كجدول تجزئة، على سبيل المثال (تعتمد طريقة التنفيذ الفعلية على بيئة العمل التي يعمل فيها). طريقة عمل جدول التجزئة هي أنه يأخذ مفتاحًا كمدخل ويُجزئ هذا المفتاح بطريقة حتمية إلى قيمة رقمية. تُستخدم هذه القيمة الرقمية كمفتاح فعلي يُخزن بناءً عليه القيمة المرتبطة. ثم، إذا حاولت الوصول إلى نفس المفتاح مرة أخرى، ستُعالج دالة التجزئة المفتاح وتُرجع نفس النتيجة الرقمية، والتي تُستخدم بعد ذلك للبحث عن القيمة المرتبطة. هذا يوفر وقت بحث فعال جدًا بمعدل O(1) في المتوسط.

يمكن تنفيذ جداول التجزئة كمصفوفات مع دوال تجزئة تُنتج فهارس مصفوفة ضمن نطاق محدد. في هذه الطريقة، يكون اختيار حجم المصفوفة مهمًا، وكذلك دالة التجزئة. على سبيل المثال، ماذا لو أنتجت دالة التجزئة نفس القيمة لمفتاحين مختلفين؟ يُسمى هذا تصادمًا. إحدى طرق التعامل مع التصادمات هي تخزين كلا زوجي المفتاح-القيمة في ذلك الفهرس. ثم، عند البحث عن أي منهما، يجب التكرار خلال دلو العناصر للعثور على المفتاح الذي تبحث عنه. دالة تجزئة جيدة تقلل التصادمات للحفاظ على وقت بحث فعال.

هنا، لن نهتم بتفاصيل التجزئة أو تنفيذ جدول التجزئة، بل سنحاول فقط فهم عام لكيفية عملها.

# --instructions--

لننشئ الوظائف الأساسية لجدول التجزئة. أنشأنا لك دالة تجزئة بسيطة للاستخدام. يمكنك تمرير قيمة سلسلة نصية إلى الدالة `hash` وستُرجع قيمة مجزأة يمكنك استخدامها كمفتاح للتخزين. خزّن العناصر بناءً على هذه القيمة المجزأة في الكائن `this.collection`. أنشئ هذه الطرق الثلاث: `add`، و`remove`، و`lookup`. الأولى يجب أن تقبل زوج مفتاح-قيمة للإضافة إلى جدول التجزئة. الثانية يجب أن تزيل زوج المفتاح-القيمة عند تمرير مفتاح. الثالثة يجب أن تقبل مفتاحًا وتُرجع القيمة المرتبطة أو `null` إذا لم يكن المفتاح موجودًا.

تأكد من كتابة شفرتك مع الأخذ في الاعتبار التصادمات!

**ملاحظة:** لن تنجح اختبارات طريقة `remove` حتى تُنفذ طريقتي `add` و`lookup` بشكل صحيح.

# --hints--

يجب أن تكون دالة `hash` صالحة.

```js
let calls = 0;
const ourHash = string => {
  calls++;
  let hashCode = 0;
  for (let i = 0; i < string.length; i++) {
    hashCode += string.charCodeAt(i);
  }
  return hashCode;
};

assert.strictEqual(hash('yek'),ourHash('yek'));

assert.strictEqual(hash('key'),ourHash('key'));
assert.strictEqual(hash('key1'),ourHash('key1'));
assert.strictEqual(hash('key2'),ourHash('key2'));
assert.strictEqual(hash('key3'),ourHash('key3'));

assert.strictEqual(hash('1key'),ourHash('1key'));
assert.strictEqual(hash('ke1y'),ourHash('ke1y'));
assert.strictEqual(hash('altKey'),ourHash('altKey'));

assert.strictEqual(called,calls); 
```

يجب أن يوجد هيكل بيانات `HashTable`.

```js
assert.isDefined(HashTable); 
let newHashTable = new HashTable(); 
assert.isObject(newHashTable); 
```

يجب أن يحتوي `HashTable` على طريقة `add`.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.add); 
```

يجب أن يحتوي `HashTable` على طريقة `lookup`.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.lookup); 
```

يجب أن يحتوي `HashTable` على طريقة `remove`.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.remove); 
```

يجب أن تضيف طريقة `add` أزواج المفتاح-القيمة ويجب أن تُرجع طريقة `lookup` القيم المرتبطة بمفتاح معين.

```js
let newHashTable = new HashTable(); 
newHashTable.add('key','value');
assert.strictEqual(newHashTable.lookup('key'),'value'); 
```

يجب أن تقبل طريقة `remove` مفتاحًا كمدخل ويجب أن تزيل زوج المفتاح-القيمة المرتبط.

```js
let hashValue = hash('key');
let newHashTable = new HashTable(); 
newHashTable.add('key','value');
newHashTable.remove('key'); 
assert.notProperty(newHashTable.collection,hashValue); 
```

يجب أن تزيل طريقة `remove` زوج المفتاح-القيمة الصحيح فقط.

```js
let newHashTable = new HashTable(); 
let hashValue = hash('key');
newHashTable.add('key','value');
newHashTable.add('yek','value');
newHashTable.add('altKey','value');
newHashTable.remove('yek');

assert.notExists(newHashTable.lookup('yek'));
assert.exists(newHashTable.lookup('altKey'));
assert.exists(newHashTable.lookup('key'));

newHashTable.remove('key');

assert.notProperty(newHashTable.collection,hashValue); 
assert.exists(newHashTable.lookup('altKey'));

```

يجب إضافة العناصر باستخدام دالة التجزئة.

```js
let newHashTable = new HashTable(); 
called = 0;
newHashTable.add('key1', 'value1')
newHashTable.add('key2', 'value2');
newHashTable.add('key3', 'value3');
assert.strictEqual(called,3); 
```

يجب أن يتعامل جدول التجزئة مع التصادمات.

```js
let newHashTable = new HashTable(); 
called = 0;
newHashTable.add('key1','value1');
newHashTable.add('1key', 'value2'); 
newHashTable.add('ke1y', 'value3'); 
assert.strictEqual(newHashTable.lookup('key1'),'value1');
assert.strictEqual(newHashTable.lookup('1key'),'value2');
assert.strictEqual(newHashTable.lookup('ke1y'),'value3');
```

# --seed--

## --seed-contents--

```js
var called = 0;
var hash = string => {
  called++;
  var hashed = 0;
  for (var i = 0; i < string.length; i++) {
    hashed += string.charCodeAt(i);
  }
  return hashed;
};
var HashTable = function() {
  this.collection = {};
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
var called = 0;
var hash = (string) => {
  called++;
  var hash = 0;
  for (var i = 0; i < string.length; i++) { hash += string.charCodeAt(i); }
  return hash;
};
var HashTable = function() {
  this.collection = {};
  // Only change code below this line

  this.add = function(key, val) {
    var theHash = hash(key);
    if (!this.collection.hasOwnProperty(theHash)) {
      this.collection[theHash] = {};
    }
    this.collection[theHash][key] = val;
  }

  this.remove = function(key) {
    var theHash = hash(key);
    var hashedObj = this.collection[theHash];
    if (hashedObj.hasOwnProperty(key)) {
      delete hashedObj[key];
    }
    if (!Object.keys(hashedObj).length) {
      delete this.collection[theHash];
    }
  }

  this.lookup = function(key) {
    var theHash = hash(key);
    if (this.collection.hasOwnProperty(theHash)) {
      return this.collection[theHash][key];
    }
    return null
  }
  // Only change code above this line
};
```
