---
id: 587d825a367417b2b2512c87
title: إنشاء قائمة مرتبطة مزدوجة
challengeType: 1
forumTopicId: 301626
dashedName: create-a-doubly-linked-list
---

# --description--

كل قوائم المرتبطة التي أنشأناها حتى الآن هي قوائم مرتبطة أحادية. هنا، سننشئ <dfn>قائمة مرتبطة مزدوجة</dfn>. كما يوحي الاسم، تحتوي العقد في القائمة المرتبطة المزدوجة على مراجع للعقدة التالية والسابقة في القائمة.

هذا يسمح لنا بالتنقل في القائمة في كلا الاتجاهين، لكنه يتطلب أيضًا استخدام ذاكرة أكثر لأن كل عقدة يجب أن تحتوي على مرجع إضافي للعقدة السابقة في القائمة.

# --instructions--

لقد وفرنا كائن `Node` وبدأنا في إنشاء `DoublyLinkedList`. لنضف طريقتين إلى القائمة المرتبطة المزدوجة تسمى `add` و `remove`. يجب أن تضيف طريقة `add` العنصر المعطى إلى القائمة، بينما يجب أن تزيل طريقة `remove` كل التكرارات لعنصر معين في القائمة.

كن حذرًا في التعامل مع أي حالات حافة محتملة عند كتابة هذه الطرق، مثل الحذف لأول أو لآخر عنصر. أيضًا، يجب أن تُرجع إزالة أي عنصر من قائمة فارغة القيمة `null`.

# --hints--

يجب أن يوجد هيكل البيانات `DoublyLinkedList`.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return typeof test == 'object';
  })()
);
```

يجب أن تحتوي `DoublyLinkedList` على طريقة تسمى `add`.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.add == undefined) {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

يجب أن تحتوي `DoublyLinkedList` على طريقة تسمى `remove`.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.remove == undefined) {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

يجب أن تُرجع إزالة عنصر من قائمة فارغة القيمة `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return test.remove(100) == null;
  })()
);
```

يجب أن تضيف طريقة `add` عناصر إلى القائمة.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(5);
    test.add(6);
    test.add(723);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '56723';
  })()
);
```

يجب أن تحتفظ كل عقدة بمرجع للعقدة السابقة.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(50);
    test.add(68);
    test.add(73);
    var result = [];
    var curr = test.tail;
    while (curr) {
      result.push(curr.data);
      curr = curr.prev;
    }
    return result.join('') == '736850';
  })()
);
```

يجب أن يكون من الممكن إزالة العنصر الأول من القائمة.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(25);
    test.add(35);
    test.add(60);
    test.remove(25);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '3560';
  })()
);
```

يجب أن يكون من الممكن إزالة العنصر الأخير من القائمة.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(25);
    test.add(35);
    test.add(60);
    test.remove(60);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '2535';
  })()
);
```


# --seed--

## --seed-contents--

```js
var Node = function(data, prev) {
  this.data = data;
  this.prev = prev;
  this.next = null;
};
var DoublyLinkedList = function() {
  this.head = null;
  this.tail = null;
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
