---
id: 587d825d367417b2b2512c96
title: البحث بالعمق أولاً
challengeType: 1
forumTopicId: 301640
dashedName: depth-first-search
---

# --description--

مماثل لـ <dfn>البحث بالاتساع أولاً</dfn>، سنتعلم هنا عن خوارزمية أخرى لتصفح الرسم البياني تُسمى <dfn>البحث بالعمق أولاً</dfn>.

بينما يبحث البحث بالاتساع أولاً على أطوال الحواف المتزايدة بعيدًا عن العقدة المصدر، يذهب <dfn>البحث بالعمق أولاً</dfn> أولاً في مسار من الحواف لأبعد حد ممكن.

بمجرد أن يصل إلى نهاية مسار، يعود البحث إلى آخر عقدة بها مسار حواف غير مُزار ويواصل البحث.

تُظهر الرسوم المتحركة أدناه كيف تعمل الخوارزمية. تبدأ الخوارزمية بالعقدة العليا وتزور العقد حسب الترتيب الرقمي.

<img alt="animation showing the depth first search algorithm" src='https://cdn.freecodecamp.org/curriculum/coding-interview-prep/depth-first-search.gif'>

لاحظ كيف أنه، على عكس البحث بالاتساع أولاً، لا يزور كل عقدة جميع جيرانها عند زيارتها. بل يزور أولاً أحد جيرانها ويستمر في ذلك المسار حتى لا تبقى عقد لزيارتها في ذلك المسار.

لتنفيذ هذه الخوارزمية، ستحتاج إلى استخدام مكدس. المكدس هو مصفوفة يُزال منها العنصر الأخير المُضاف أولاً. يُعرف هذا أيضًا باسم هيكل البيانات <dfn>Last-In-First-Out</dfn>. المكدس مفيد في خوارزميات البحث بالعمق أولاً لأننا عند إضافة الجيران إلى المكدس، نريد زيارة الجيران المضافين حديثًا أولاً وإزالتهم من المكدس.

الناتج البسيط لهذه الخوارزمية هو قائمة بالعقد التي يمكن الوصول إليها من عقدة معينة. لذلك، ستحتاج أيضًا إلى تتبع العقد التي تزورها.

# --instructions--

اكتب دالة `dfs()` تأخذ مصفوفة مجاورة غير موجهة `graph` ومسمى عقدة `root` كمعلمات. سيكون مسمى العقدة هو القيمة الرقمية للعقدة بين `0` و `n - 1`، حيث `n` هو العدد الكلي للعقد في الرسم البياني.

يجب أن تُخرج دالتك مصفوفة بكل العقد التي يمكن الوصول إليها من `root`.

# --hints--

الرسم البياني المدخل `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` مع عقدة بداية `1` يجب أن يُرجع مصفوفة تحتوي على `0`، `1`، `2`، و `3`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 1);
  })(),
  [0, 1, 2, 3]
);
```

الرسم البياني المدخل `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` مع عقدة بداية `3` يجب أن يُرجع مصفوفة تحتوي على `3`، `2`، `1`، و `0`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })(),
  [3, 2, 1, 0]
);
```

الرسم البياني المدخل `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` مع عقدة بداية `1` يجب أن يُرجع مصفوفة بها أربعة عناصر.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 1);
  })().length === 4
);
```

الرسم البياني المدخل `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` مع عقدة بداية `3` يجب أن يُرجع مصفوفة تحتوي على `3`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 0],
      [0, 0, 0, 0]
    ];
    return dfs(graph, 3);
  })(),
  [3]
);
```

الرسم البياني المدخل `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` مع عقدة بداية `3` يجب أن يُرجع مصفوفة بها عنصر واحد.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 0],
      [0, 0, 0, 0]
    ];
    return dfs(graph, 3);
  })().length === 1
);
```

الرسم البياني المدخل `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` مع عقدة بداية `3` يجب أن يُرجع مصفوفة تحتوي على `2` و `3`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })(),
  [2, 3]
);
```

الرسم البياني المدخل `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` مع عقدة بداية `3` يجب أن يُرجع مصفوفة بها عنصران.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })().length === 2
);
```

الرسم البياني المدخل `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` مع عقدة بداية `0` يجب أن يُرجع مصفوفة تحتوي على `0` و `1`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 0);
  })(),
  [0, 1]
);
```

الرسم البياني المدخل `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` مع عقدة بداية `0` يجب أن يُرجع مصفوفة بها عنصران.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 0);
  })().length === 2
);
```

# --seed--

## --seed-contents--

```js
function dfs(graph, root) {

}

var exDFSGraph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
console.log(dfs(exDFSGraph, 3));
```

# --solutions--

```js
function dfs(graph, root) {
    var stack = [];
    var tempV;
    var visited = [];
    var tempVNeighbors = [];
    stack.push(root);
    while (stack.length > 0) {
        tempV = stack.pop();
        if (visited.indexOf(tempV) == -1) {
            visited.push(tempV);
            tempVNeighbors = graph[tempV];
            for (var i = 0; i < tempVNeighbors.length; i++) {
                if (tempVNeighbors[i] == 1) {
                    stack.push(i);
                }
            }
        }
    }
    return visited;
}
```
