---
id: 587d8251367417b2b2512c62
title: إنشاء صنف قائمة مرتبطة
challengeType: 1
forumTopicId: 301628
dashedName: create-a-linked-list-class
---

# --description--

لننشئ صنف `linked list`. يجب أن تبدأ كل قائمة مرتبطة ببعض الخصائص الأساسية: `head` (العنصر الأول في قائمتك) و`length` (عدد العناصر في قائمتك). أحيانًا سترى تطبيقات للقوائم المرتبطة تتضمن `tail` للعنصر الأخير في القائمة، لكننا سنكتفي الآن بهذين الخاصيتين فقط. في كل مرة نضيف عنصرًا إلى القائمة المرتبطة، يجب زيادة خاصية `length` بمقدار واحد.

نحتاج إلى طريقة لإضافة عناصر إلى القائمة المرتبطة، لذلك أول طريقة سننشئها هي طريقة `add`.

إذا كانت قائمتنا فارغة، فإن إضافة عنصر إلى القائمة المرتبطة أمر بسيط: نغلف هذا العنصر في صنف `Node`، ثم نُعيّن تلك العقدة إلى `head` في القائمة المرتبطة.

لكن ماذا لو كانت القائمة تحتوي بالفعل على عضو واحد أو أكثر؟ كيف نضيف عنصرًا إلى القائمة؟ تذكر أن كل عقدة في القائمة المرتبطة لها خاصية `next`. لإضافة عقدة إلى القائمة، ابحث عن العقدة الأخيرة في القائمة، ووجّه خاصية `next` لتلك العقدة الأخيرة إلى العقدة الجديدة. (تلميح: تعرف أنك وصلت إلى نهاية القائمة المرتبطة عندما تكون خاصية `next` للعقدة تساوي `null`.)

# --instructions--

اكتب طريقة add تُعيّن أول عقدة تضيفها إلى القائمة المرتبطة إلى `head`؛ بعد ذلك، عند إضافة عقدة، يجب أن تشير كل عقدة إلى العقدة الجديدة عبر خاصية `next` للعقدة السابقة.

ملاحظة

يجب أن تزيد خاصية `length` في قائمتك بمقدار واحد في كل مرة تضيف فيها عنصرًا إلى القائمة المرتبطة.

# --hints--

يجب أن يحتوي صنف `LinkedList` على طريقة `add`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.add === 'function';
  })()
);
```

يجب أن يُعيّن صنف `LinkedList` خاصية `head` إلى أول عقدة تُضاف.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    return test.head().element === 'cat';
  })()
);
```

يجب أن يحتوي `node` السابق في صنف `LinkedList` على مرجع لأحدث عقدة تم إنشاؤها.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('fish');
    return test.head().next.element === 'dog' && test.head().next.next.element === 'fish';
  })()
);
```

يجب أن يساوي `size` في صنف `LinkedList` عدد العقد في القائمة المرتبطة.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    return test.size() === 2;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.head = function(){
    return head;
  };

  this.size = function(){
    return length;
  };

  this.add = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() { 
  var length = 0; 
  var head = null; 

  var Node = function(element){
    this.element = element; 
    this.next = null; 
  }; 

  this.head = function(){
    return head;
  };

  this.size = function(){
    return length;
  };

  this.add = function(element){
    // Only change code below this line
    if (head == null) {
      head = new Node(element);
    } 
    else {
      let currentNode = head;
      while (currentNode.next != null) {
        // currentNode.next will be last node of linked list after loop
        currentNode = currentNode.next;
      }
      currentNode.next = new Node(element);
    }
    length++;
    // Only change code above this line
  };
}
```
