---
id: 587d8258367417b2b2512c80
title: حذف عقدة ورقية في شجرة البحث الثنائي
challengeType: 1
forumTopicId: 301637
dashedName: delete-a-leaf-node-in-a-binary-search-tree
---

# --description--

هذا هو التحدٍّ الأول من ثلاثة حيث سننفذ عملية أكثر صعوبة في أشجار البحث الثنائي: الحذف. الحذف صعب لأن إزالة العقد تكسر الروابط في الشجرة. يجب إعادة إنشاء هذه الروابط بعناية لضمان الحفاظ على هيكل الشجرة الثنائية. لبعض عمليات الحذف، يعني هذا أنه يجب إعادة ترتيب الشجرة. بشكل عام، ستواجه أحد ثلاثة حالات عند محاولة حذف عقدة: عقدة ورقية: الهدف للحذف لا يملك أطفالًا. طفل واحد: الهدف للحذف يملك طفلًا واحدًا فقط. طفلان: الهدف للحذف يملك عقدتين فرعيتين. إزالة عقدة ورقية سهلة، فقط نحذفها. حذف عقدة بطفل واحد أيضًا سهل نسبيًا، نحذفها ونربط والدها بطفل العقدة التي حذفناها. حذف عقدة بطفلين أصعب، لأنه يخلق عقدتين فرعيتين تحتاجان لإعادة ربطهما بالشجرة الأصلية. سنرى كيفية التعامل مع هذه الحالة في التحدٍّ الثالث. بالإضافة لذلك، يجب أن تنتبه لبعض الحالات الخاصة عند التعامل مع الحذف. ماذا لو كانت الشجرة فارغة؟ ماذا لو كانت العقدة للحذف هي العقدة الجذرية؟ ماذا لو كانت الشجرة تحتوي على عنصرين فقط؟ الآن، دعنا نتعامل مع الحالة الأولى حيث نحذف عقدة ورقية.

# --instructions--

أنشئ طريقة على شجرتنا الثنائية تُسمى `remove`. سنبني منطق عملية الحذف هنا. أولًا، ستحتاج إلى إنشاء دالة داخل remove تجد العقدة التي نحاول حذفها في الشجرة الحالية. إذا لم تكن العقدة موجودة في الشجرة، فيجب أن تُعيد `remove` القيمة `null`. الآن، إذا كانت العقدة المستهدفة عقدة ورقية بلا أبناء، فيجب ضبط مرجع الأصل إليها على `null`. هذا يحذف العقدة من الشجرة فعليًا. وللقيام بذلك، سيتعين عليك تتبّع أصل العقدة التي نحاول حذفها أيضًا. وسيكون من المفيد أيضًا إنشاء وسيلة لتتبّع عدد الأبناء الذين تملكهم العقدة المستهدفة، إذ سيحدد ذلك الحالة التي يندرج تحتها حذفنا. سنتناول الحالتين الثانية والثالثة في التحديات القادمة. حظًا موفقًا!

# --hints--

يجب أن يوجد هيكل البيانات `BinarySearchTree`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

يجب أن تحتوي شجرة البحث الثنائي على طريقة تسمى `remove`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

محاولة حذف عنصر من شجرة فارغة يجب أن تُرجع `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    return test.remove(100) == null;
  })()
);
```

محاولة حذف عنصر غير موجود يجب أن تُرجع `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(15);
    test.add(30);
    return test.remove(100) == null;
  })()
);
```

إذا لم يكن للعقدة الجذرية أطفال، يجب أن يؤدي حذفها إلى تعيين الجذر إلى `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(500);
    test.remove(500);
    return test.inorder() == null;
  })()
);
```

يجب أن تزيل طريقة `remove` العقد الورقية من الشجرة.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(5);
    test.add(3);
    test.add(7);
    test.add(6);
    test.add(10);
    test.add(12);
    test.remove(3);
    test.remove(12);
    test.remove(10);
    return test.inorder().join('') == '567';
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    },
    inorder: function() {
      if (this.root == null) {
        return null;
      } else {
        var result = new Array();
        function traverseInOrder(node) {
          if (node.left != null) {
            traverseInOrder(node.left);
          }
          result.push(node.value);
          if (node.right != null) {
            traverseInOrder(node.right);
          }
        }
        traverseInOrder(this.root);
        return result;
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}

function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
}
```

# --solutions--

```js
// solution required
```
