---
id: 587d8251367417b2b2512c65
title: إزالة عناصر من قائمة مرتبطة حسب الفهرس
challengeType: 1
forumTopicId: 301711
dashedName: remove-elements-from-a-linked-list-by-index
---

# --description--

قبل أن ننتقل إلى هيكل بيانات آخر، دعنا نمارس بعض التمارين الأخيرة مع القوائم المرتبطة.

لنكتب طريقة `removeAt` التي تزيل الـ `element` عند `index` معين. يجب أن تُسمى الطريقة `removeAt(index)`. لإزالة `element` عند `index` معين، نحتاج إلى تتبع عدد كل عقدة أثناء التنقل في القائمة المرتبطة.

تقنية شائعة تُستخدم في التكرار خلال عناصر القائمة المرتبطة تعتمد على <dfn>"العداء"</dfn>، أو الحارس، الذي "يشير" إلى العقد التي يقارنها شفرتك. في حالتنا، نبدأ من `head` القائمة، ونبدأ بمتغير `currentIndex` يبدأ من `0`. يجب أن يزيد `currentIndex` بمقدار واحد لكل عقدة نتجاوزها.

تمامًا مثل طريقة `remove(element)` التي <a href="/learn/coding-interview-prep/data-structures/remove-elements-from-a-linked-list" target="_blank" rel="noopener noreferrer nofollow">تطرقنا إليها في درس سابق</a>، يجب أن نحرص على عدم ترك بقية القائمة يتيمة عند إزالة العقدة في طريقة `removeAt(index)`. نحافظ على تتابع العقد بالتأكد من أن العقدة التي تشير إلى العقدة المحذوفة تشير إلى العقدة التالية.

# --instructions--

اكتب طريقة `removeAt(index)` تزيل وتُرجع عقدة عند `index` معين. يجب أن تُرجع الطريقة `null` إذا كان `index` المعطى سالبًا أو أكبر من أو يساوي `length` للقائمة المرتبطة.

**ملاحظة:** تذكر تتبع عدد `currentIndex`.

# --hints--

يجب أن يحتوي صنف `LinkedList` على طريقة `removeAt`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.removeAt === 'function';
  })()
);
```

يجب أن تُقلل طريقة `removeAt` من `length` القائمة المرتبطة بمقدار واحد.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.removeAt(1);
    return test.size() === 2;
  })()
);
```

يجب أن تزيل طريقة `removeAt` العنصر عند الفهرس المحدد من القائمة المرتبطة.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.add('bird');
    test.removeAt(1);
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"kitten","next":{"element":"bird","next":null}}}'
    );
  })()
);
```

عندما يكون هناك عنصر واحد فقط في القائمة المرتبطة، يجب أن تزيل طريقة `removeAt` العنصر عند الفهرس المحدد وتُرجعه، وتقلل طول القائمة المرتبطة.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    var removedItem = test.removeAt(0);
    return test.head() === null && test.size() === 0 && removedItem === 'cat';
  })()
);
```

يجب أن تُرجع طريقة `removeAt` العنصر الخاص بالعقدة المحذوفة.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.removeAt(1) === 'dog';
  })()
);
```

يجب أن تُرجع طريقة `removeAt` `null` ولا تتغير القائمة المرتبطة إذا كان الفهرس المعطى أقل من `0`.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    var removedItem = test.removeAt(-1);
    return (
      removedItem === null &&
      JSON.stringify(test.head()) ===
        '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

يجب أن تُرجع طريقة `removeAt` القيمة `null` وألا تتغير القائمة المرتبطة إذا كان الفهرس المعطى أكبر من أو يساوي `length` للقائمة.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    var removedItem = test.removeAt(3);
    return (
      removedItem === null &&
      JSON.stringify(test.head()) ===
        '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
      head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function (element) {
    this.element = element;
    this.next = null;
  };

  this.size = function () {
    return length;
  };

  this.head = function () {
    return head;
  };

  this.add = function (element) {
    var node = new Node(element);
    if (head === null) {
      head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.removeAt = function (index) {
    var currentNode = head;
    var previous = head;
    var count = 0;
    if (index >= length || index < 0) {
      return null;
    }
    if (index === 0) {
      var removed = head.element;
      head = currentNode.next;
    } else {
      while (count < index) {
        previous = currentNode;
        currentNode = currentNode.next;
        count++;
      }
      var removed = previous.next.element;
      previous.next = currentNode.next;
    }
    length--;
    return removed;
  };
}
```
