---
id: 587d8251367417b2b2512c61
title: العمل مع العقد في قائمة مرتبطة
challengeType: 1
forumTopicId: 301721
dashedName: work-with-nodes-in-a-linked-list
---

# --description--

هيكل البيانات الشائع الآخر الذي ستصادفه في علوم الحاسوب هو <dfn>القائمة المرتبطة</dfn>. القائمة المرتبطة هي مجموعة خطية من عناصر البيانات، تُسمى "عقدًا"، كل منها يشير إلى العقدة التالية. تحتوي كل <dfn>عقدة</dfn> في القائمة المرتبطة على قطعتين رئيسيتين من المعلومات: `element` نفسه، ومرجع إلى `node` التالية.

تخيل أنك في صف رقصة الكونغا. يديك على الشخص التالي في الصف، والشخص خلفك يده على يدك. يمكنك رؤية الشخص أمامك مباشرة، لكنه يحجب رؤية الأشخاص الآخرين في الصف. العقدة تشبه الشخص في صف الكونغا: يعرف من هو ويمكنه رؤية الشخص التالي فقط، لكنه لا يدرك الأشخاص الآخرين أمامه أو خلفه.

# --instructions--

في محرر الشفرة الخاص بنا، أنشأنا عقدتين، `Kitten` و `Puppy`، وربطنا يدويًا عقدة `Kitten` بعقدة `Puppy`.

أنشئ عقدتي `Cat` و `Dog` وأضفهما يدويًا إلى الصف.

# --hints--

يجب أن تحتوي عقدة `Puppy` على مرجع إلى عقدة `Cat`.

```js
assert(Puppy.next.element === 'Cat');
```

يجب أن تحتوي عقدة `Cat` على مرجع إلى عقدة `Dog`.

```js
assert(Cat.next.element === 'Dog');
```

# --seed--

## --seed-contents--

```js
var Node = function(element) {
  this.element = element;
  this.next = null;
};
var Kitten = new Node('Kitten');
var Puppy = new Node('Puppy');

Kitten.next = Puppy;
// Only change code below this line
```

# --solutions--

```js
// solution required
```
