---
id: 587d825b367417b2b2512c8d
title: إنشاء خريطة JavaScript ES6
challengeType: 1
forumTopicId: 301635
dashedName: create-an-es6-javascript-map
---

# --description--

يوفر لنا الإصدار الجديد من JavaScript كائن Map مدمج يوفّر الكثير من الوظائف التي كتبناها يدويًا في التحدي السابق. هذا الكائن Map، رغم تشابهه مع كائنات JavaScript العادية، يقدم بعض الوظائف المفيدة التي تفتقر إليها الكائنات العادية. على سبيل المثال، تتعقب خريطة ES6 ترتيب الإدخال للعناصر التي تُضاف إليها. إليك نظرة أكثر شمولًا على طرقه:

- `.has(key)` تُرجع true أو false بناءً على وجود المفتاح
- `.get(key)` تُرجع القيمة المرتبطة بالمفتاح
- `.set(key, value)` تُعيّن زوج مفتاح وقيمة جديد
- `.delete(key)` تحذف زوج المفتاح والقيمة
- `.clear()` تمسح كل أزواج المفتاح والقيمة
- `.keys()` تُرجع كائن مكرر جديد للخريطة يحتوي على كل المفاتيح بترتيب الإدخال
- `.values()` تُرجع كائن مكرر جديد للخريطة يحتوي على كل القيم بترتيب الإدخال
- `.entries()` تُرجع كائن مكرر جديد للخريطة يحتوي على كل أزواج المفتاح والقيمة بترتيب الإدخال

# --instructions--

عرّف كائن Map في JavaScript وعيّنه إلى متغير اسمه myMap. أضف زوج المفتاح والقيمة `freeCodeCamp`، `Awesome!` إليه.

# --hints--

يجب أن يوجد كائن `myMap`.

```js
assert(typeof myMap === 'object');
```

يجب أن يحتوي `myMap` على زوج المفتاح والقيمة `freeCodeCamp`، `Awesome!`.

```js
assert(myMap.get('freeCodeCamp') === 'Awesome!');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const myMap = new Map();

myMap.set("freeCodeCamp", "Awesome!");
```
