---
id: 587d8254367417b2b2512c71
title: حذف عناصر من مجموعة في ES6
challengeType: 1
forumTopicId: 301713
dashedName: remove-items-from-a-set-in-es6
---

# --description--

لنتمرن على حذف عناصر من مجموعة ES6 باستخدام الطريقة `delete`.

أولًا، أنشئ مجموعة ES6:

```js
var set = new Set([1,2,3]);
```

الآن احذف عنصرًا من مجموعتك باستخدام الطريقة `delete`.

```js
set.delete(1);
console.log([...set]) // should return [ 2, 3 ]
```

# --instructions--

الآن، أنشئ مجموعة تحتوي على الأعداد الصحيحة 1، 2، 3، 4، و5.

احذف القيم 2 و5، ثم أعد المجموعة.

# --hints--

يجب أن تحتوي مجموعتك على القيم 1، 3، و4

```js
assert(
  (function () {
    var test = checkSet();
    return test.has(1) && test.has(3) && test.has(4) && test.size === 3;
  })()
);
```

# --seed--

## --seed-contents--

```js
function checkSet(){
  // Only change code below this line
  var set = null;

  // Only change code above this line
  return set;   
}
```

# --solutions--

```js
function checkSet(){
var set = new Set([1,2,3,4,5]);
set.delete(2);
set.delete(5);
return set;}
```
