---
id: 587d8250367417b2b2512c5e
title: تعلّم كيف يعمل المكدس
challengeType: 1
forumTopicId: 301705
dashedName: learn-how-a-stack-works
---

# --description--

من المحتمل أنك معتاد على مكدس الكتب على طاولتك. ربما استخدمت ميزة التراجع في محرر النصوص. وربما اعتدت أيضًا على الضغط على زر الرجوع في هاتفك للعودة إلى العرض السابق في التطبيق.

هل تعرف ما المشترك بينهم جميعًا؟ كلهم يخزنون البيانات بطريقة تسمح لك بالتنقل إلى الخلف.

أعلى كتاب في المكدس هو الذي وُضع هناك أخيرًا. إذا أزلت ذلك الكتاب من أعلى المكدس، ستُتاح لك رؤية الكتاب الذي وُضع قبله وهكذا.

إذا فكرت في الأمر، في كل الأمثلة السابقة، تحصل على خدمة من نوع <dfn>Last-In-First-Out</dfn>. سنحاول تقليد هذا في شفرتنا.

تسمى هذه الطريقة في تخزين البيانات <dfn>مكدس</dfn>. على وجه الخصوص، يجب علينا تنفيذ الطريقة `push()` التي تضيف كائنات JavaScript إلى أعلى المكدس؛ والطريقة `pop()` التي تزيل كائن JavaScript الموجود في أعلى المكدس في اللحظة الحالية.

# --instructions--

هنا لدينا مكدس من واجبات منزلية ممثل كمصفوفة: `"BIO12"` في القاعدة، و`"PSY44"` في أعلى المكدس.

عدّل المصفوفة المعطاة وتعامل معها كـ `stack` باستخدام طرق JavaScript المذكورة أعلاه. أزل العنصر الأعلى `"PSY44"` من الـ stack. ثم أضف `"CS50"` ليكون العنصر الأعلى الجديد في الـ stack.

# --hints--

يجب أن يحتوي `homeworkStack` على 4 عناصر فقط.

```js
assert(homeworkStack.length === 4);
```

يجب أن يكون العنصر الأخير في `homeworkStack` هو `"CS50"`.

```js
assert(homeworkStack[3] === 'CS50');
```

يجب ألا يحتوي `homeworkStack` على `"PSY44"`.

```js
assert(homeworkStack.indexOf('PSY44') === -1);
```

يجب عدم تغيير الإعلان الأولي لـ `homeworkStack`.

```js
assert(
  __helpers.removeJSComments(code).match(/=/g).length === 1 &&
    /homeworkStack\s*=\s*\["BIO12"\s*,\s*"HIS80"\s*,\s*"MAT122"\s*,\s*"PSY44"\]/.test(
      __helpers.removeJSComments(code)
    )
);
```

# --seed--

## --seed-contents--

```js
var homeworkStack = ["BIO12","HIS80","MAT122","PSY44"];
// Only change code below this line
```

# --solutions--

```js
// solution required
```
