---
id: 587d8258367417b2b2512c7f
title: استخدام البحث بالاتساع أولاً في شجرة البحث الثنائي
challengeType: 1
forumTopicId: 301718
dashedName: use-breadth-first-search-in-a-binary-search-tree
---

# --description--

هنا سنُقدّم طريقة أخرى لاستعراض الشجرة: البحث بالاتساع أولاً. على عكس طرق البحث بالعمق أولاً من التحدي السابق، يستعرض البحث بالاتساع أولاً كل العقد في مستوى معين داخل الشجرة قبل الانتقال إلى المستوى التالي. عادةً ما تُستخدم قوائم الانتظار كـ هياكل بيانات مساعدة في تصميم خوارزميات البحث بالاتساع أولاً.

في هذه الطريقة، نبدأ بإضافة العقدة الجذرية إلى قائمة انتظار. ثم نبدأ حلقة حيث نُخرج أول عنصر من قائمة الانتظار، نضيفه إلى مصفوفة جديدة، ثم نفحص كلتا الشجرتين الفرعيتين له. إذا لم تكن أطفاله فارغة، تُضاف كل واحدة منها إلى قائمة الانتظار. تستمر هذه العملية حتى تصبح قائمة الانتظار فارغة.

# --instructions--

لننشئ طريقة بحث بالاتساع أولاً في شجرتنا تُسمى `levelOrder`. يجب أن تُرجع هذه الطريقة مصفوفة تحتوي على قيم كل عقد الشجرة، مستعرضة بطريقة البحث بالاتساع أولاً. تأكد من إرجاع القيم في المصفوفة، وليس العقد نفسها. يجب استعراض المستوى من اليسار إلى اليمين. بعد ذلك، لنكتب طريقة مشابهة تُسمى `reverseLevelOrder` تؤدي نفس البحث ولكن في الاتجاه العكسي (من اليمين إلى اليسار) في كل مستوى.

# --hints--

يجب أن يوجد هيكل البيانات `BinarySearchTree`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

يجب أن تحتوي شجرة البحث الثنائي على طريقة تُسمى `levelOrder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.levelOrder == 'function';
  })()
);
```

يجب أن تحتوي شجرة البحث الثنائي على طريقة تُسمى `reverseLevelOrder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.reverseLevelOrder == 'function';
  })()
);
```

يجب أن تُرجع طريقة `levelOrder` مصفوفة من قيم عقد الشجرة التي تم استعراضها بترتيب المستويات.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.levelOrder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.levelOrder().join('') == '719038102546';
  })()
);
```

يجب أن تُرجع طريقة `reverseLevelOrder` مصفوفة من قيم عقد الشجرة التي تم استعراضها بترتيب المستويات العكسي.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.reverseLevelOrder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.reverseLevelOrder().join('') == '791108305264';
  })()
);
```

يجب أن تُرجع طريقة `levelOrder` القيمة `null` إذا كانت الشجرة فارغة.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.levelOrder !== 'function') {
      return false;
    }
    return test.levelOrder() == null;
  })()
);
```

يجب أن تُرجع طريقة `reverseLevelOrder` القيمة `null` إذا كانت الشجرة فارغة.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.reverseLevelOrder !== 'function') {
      return false;
    }
    return test.reverseLevelOrder() == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      function searchTree(node) {
        if (value < node.value) {
          if (node.left == null) {
            node.left = new Node(value);
            return;
          } else if (node.left != null) {
            return searchTree(node.left);
          }
        } else if (value > node.value) {
          if (node.right == null) {
            node.right = new Node(value);
            return;
          } else if (node.right != null) {
            return searchTree(node.right);
          }
        } else {
          return null;
        }
      }
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  this.levelOrder = (root = this.root) => {
    if(!root) return null;
    let queue = [root];
    let results = [];
    while(queue.length > 0) {
      let node = queue.shift();
      results.push(node.value);
      if(node.left) queue.push(node.left);
      if(node.right) queue.push(node.right);
    }
    return results;
  }

  this.reverseLevelOrder = (root = this.root) => {
    if(!root) return null;
    let queue = [root];
    let results = [] ;
    while ( queue.length > 0) {
      let node = queue.shift();
      results.push(node.value);
      if(node.right) queue.push(node.right);
      if(node.left ) queue.push(node.left);
    }
    return results;
  }
  // Only change code above this line
}
```
