---
id: 587d8257367417b2b2512c7d
title: إيجاد الحد الأدنى والحد الأقصى لارتفاع شجرة البحث الثنائي
challengeType: 1
forumTopicId: 301641
dashedName: find-the-minimum-and-maximum-height-of-a-binary-search-tree
---

# --description--

في التحدٍّ السابق وصفنا سيناريو يمكن أن تصبح فيه الشجرة غير متوازنة. لفهم مفهوم التوازن، دعنا نلقي نظرة على خاصية أخرى للشجرة: الارتفاع. الارتفاع في الشجرة يُمثل المسافة من العقدة الجذرية إلى أي عقدة ورقية معينة. قد تختلف المسارات في هيكل شجرة متفرعة بشكل كبير في ارتفاعها، لكن لكل شجرة سيكون هناك ارتفاع أدنى وارتفاع أقصى. إذا كانت الشجرة متوازنة، فإن هذين القيمتين تختلفان بحد أقصى بمقدار واحد. هذا يعني أنه في الشجرة المتوازنة، توجد كل العقد الورقية ضمن نفس المستوى، أو إذا لم تكن ضمن نفس المستوى فهي على بعد مستوى واحد على الأكثر.

خاصية التوازن مهمة للشجر لأنها تحدد كفاءة عمليات الشجرة. كما شرحنا في التحدٍّ السابق، نواجه تعقيدًا زمنيًا في أسوأ الحالات مع الأشجار غير المتوازنة بشدة. تُستخدم الأشجار ذاتية التوازن عادةً لمعالجة هذه المشكلة في الأشجار التي تحتوي على مجموعات بيانات ديناميكية. من الأمثلة الشائعة على هذه الأشجار شجر AVL، شجر الأحمر والأسود، وشجر B. تحتوي هذه الأشجار على منطق داخلي إضافي يعيد توازن الشجرة عندما تؤدي عمليات الإدخال أو الحذف إلى حالة عدم توازن.

**ملاحظة:** خاصية مشابهة للارتفاع هي العمق، الذي يشير إلى مدى بُعد عقدة معينة عن العقدة الجذرية.

# --instructions--

اكتب طريقتين لشجرتنا الثنائية: `findMinHeight` و `findMaxHeight`. يجب أن تُرجع هذه الطرق قيمة صحيحة تمثل الحد الأدنى والحد الأقصى للارتفاع داخل شجرة ثنائية معينة، على التوالي. إذا كانت العقدة فارغة، فلنعين لها ارتفاعًا بقيمة `-1` (هذا هو الحالة الأساسية). أخيرًا، أضف طريقة ثالثة `isBalanced` تُرجع `true` أو `false` حسب ما إذا كانت الشجرة متوازنة أم لا. يمكنك استخدام الطريقتين السابقتين اللتين كتبتهما لتحديد ذلك.

# --hints--

يجب أن يوجد هيكل البيانات `BinarySearchTree`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

يجب أن تحتوي شجرة البحث الثنائي على طريقة تسمى `findMinHeight`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.findMinHeight == 'function';
  })()
);
```

يجب أن تحتوي شجرة البحث الثنائي على طريقة تسمى `findMaxHeight`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.findMaxHeight == 'function';
  })()
);
```

يجب أن تحتوي شجرة البحث الثنائي على طريقة تسمى `isBalanced`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.isBalanced == 'function';
  })()
);
```

يجب أن تُرجع طريقة `findMinHeight` الحد الأدنى لارتفاع الشجرة.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.findMinHeight !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    return test.findMinHeight() == 1;
  })()
);
```

يجب أن تُرجع طريقة `findMaxHeight` الحد الأقصى لارتفاع الشجرة.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.findMaxHeight !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    return test.findMaxHeight() == 5;
  })()
);
```

يجب أن تُرجع الشجرة الفارغة ارتفاعًا بقيمة `-1`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.findMaxHeight !== 'function') {
      return false;
    }
    return test.findMaxHeight() == -1;
  })()
);
```

يجب أن تُرجع طريقة `isBalanced` القيمة `false` إذا كانت الشجرة شجرة بحث ثنائي غير متوازنة.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isBalanced !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    return test.isBalanced() === false;
  })()
);
```

يجب أن تُرجع طريقة `isBalanced` القيمة `true` إذا كانت الشجرة شجرة بحث ثنائي متوازنة.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isBalanced !== 'function') {
      return false;
    }
    test.add(10);
    test.add(3);
    test.add(22);
    test.add(1);
    test.add(4);
    test.add(17);
    test.add(32);
    return test.isBalanced() === true;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      function searchTree(node) {
        if (value < node.value) {
          if (node.left == null) {
            node.left = new Node(value);
            return;
          } else if (node.left != null) {
            return searchTree(node.left);
          }
        } else if (value > node.value) {
          if (node.right == null) {
            node.right = new Node(value);
            return;
          } else if (node.right != null) {
            return searchTree(node.right);
          }
        } else {
          return null;
        }
      }

      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
  this.findMinHeight = function(root = this.root) {
    // empty tree.
    if (root === null) {
      return -1;
    }
    // leaf node.
    if (root.left === null && root.right === null) {
      return 0;
    }
    if (root.left === null) {
      return this.findMinHeight(root.right) + 1;
    }
    if (root.right === null) {
      return this.findMinHeight(root.left) + 1;
    }
    const lHeight = this.findMinHeight(root.left);
    const rHeight = this.findMinHeight(root.right);
    return Math.min(lHeight, rHeight) + 1;
  };
  this.findMaxHeight = function(root = this.root) {
    // empty tree.
    if (root === null) {
      return -1;
    }
    // leaf node.
    if (root.left === null && root.right === null) {
      return 0;
    }
    if (root.left === null) {
      return this.findMaxHeight(root.right) + 1;
    }
    if (root.right === null) {
      return this.findMaxHeight(root.left) + 1;
    }
    const lHeight = this.findMaxHeight(root.left);
    const rHeight = this.findMaxHeight(root.right);
    return Math.max(lHeight, rHeight) + 1;
  };
  this.isBalanced = function(root = this.root) {
    if (root === null) {
      return true;
    }

    if (root.left === null && root.right === null) {
      return true;
    }

    if (root.left === null) {
      return this.findMaxHeight(root.right) <= 0;
    }

    if (root.right === null) {
      return this.findMaxHeight(root.left) <= 0;
    }

    const lHeight = this.findMaxHeight(root.left);
    const rHeight = this.findMaxHeight(root.right);
    if (Math.abs(lHeight - rHeight) > 1) {
      return false;
    }
    return this.isBalanced(root.left) && this.isBalanced(root.right);
  };
}
```
