---
id: 587d8254367417b2b2512c70
title: إنشاء وإضافة إلى المجموعات في ES6
challengeType: 1
forumTopicId: 301636
dashedName: create-and-add-to-sets-in-es6
---

# --description--

الآن بعد أن عملت مع ES5، ستقوم بشيء مشابه في ES6. سيكون ذلك أسهل بكثير. يحتوي ES6 على هيكل بيانات مدمج `Set` لذا العديد من العمليات التي كتبتها يدويًا أصبحت متوفرة لك الآن. لنلق نظرة:

لإنشاء مجموعة جديدة فارغة:

```js
var set = new Set();
```

يمكنك إنشاء مجموعة بقيمة واحدة:

```js
var set = new Set(1);
```

يمكنك إنشاء مجموعة بمصفوفة:

```js
var set = new Set([1, 2, 3]);
```

بمجرد إنشاء مجموعة، يمكنك إضافة القيم التي تريدها باستخدام طريقة `add`:

```js
var set = new Set([1, 2, 3]);
set.add([4, 5, 6]);
```

كتذكير، المجموعة هي هيكل بيانات لا يمكن أن يحتوي على قيم مكررة:

```js
var set = new Set([1, 2, 3, 1, 2, 3]);
// set contains [1, 2, 3] only
```

# --instructions--

لهذا التمرين، أرجع مجموعة تحتوي القيم التالية: `1, 2, 3, 'Taco', 'Cat', 'Awesome'`

# --hints--

يجب أن تحتوي مجموعتك `Set` فقط على القيم `1, 2, 3, Taco, Cat, Awesome`.

```js
assert(
  (function () {
    var test = checkSet();
    return (
      test.size == 6 &&
      test.has(1) &&
      test.has(2) &&
      test.has(3) &&
      test.has('Taco') &&
      test.has('Cat') &&
      test.has('Awesome')
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function checkSet() {
  var set = new Set([1, 2, 3, 3, 2, 1, 2, 3, 1]);
  // Only change code below this line

  // Only change code above this line
  console.log(Array.from(set));
  return set;
}

checkSet();
```

# --solutions--

```js
function checkSet(){var set = new Set([1,2,3,'Taco','Cat','Awesome']);
return set;}
```
