---
id: 587d825b367417b2b2512c8c
title: تنفيذ فرز الحزمة باستخدام كومة صغرى
challengeType: 1
forumTopicId: 301643
dashedName: implement-heap-sort-with-a-min-heap
---

# --description--

الآن بعد أن أصبح بإمكاننا إضافة العناصر وإزالتها، دعنا نرى بعض التطبيقات التي يمكن استخدام الحزم فيها. تُستخدم الحزم عادةً لتنفيذ طوابير الأولوية لأنها تخزن دائمًا العنصر ذي القيمة الأكبر أو الأصغر في الموضع الأول. بالإضافة إلى ذلك، تُستخدم لتنفيذ خوارزمية فرز تُسمى فرز الحزمة. سنرى كيف نفعل ذلك هنا. يستخدم فرز الحزمة كومة صغرى، وهي عكس الحزمة القصوى. تخزن الكومة الصغرى دائمًا العنصر ذي القيمة الأقل في موضع الجذر.

يعمل فرز الحزمة بأخذ مصفوفة غير مرتبة، وإضافة كل عنصر في المصفوفة إلى كومة صغرى، ثم استخراج كل عنصر من الكومة الصغرى إلى مصفوفة جديدة. يضمن هيكل الكومة الصغرى أن تحتوي المصفوفة الجديدة على العناصر الأصلية مرتبة من الأصغر إلى الأكبر. هذه واحدة من أكثر خوارزميات الفرز كفاءة مع أداء متوسط وأسوأ حالة بمقدار O(nlog(n)).

# --instructions--

لننفذ فرز الحزمة باستخدام كومة صغرى. لا تتردد في تعديل شفرة الحزمة القصوى الخاصة بك هنا. أنشئ كائنًا `MinHeap` يحتوي على طرق `insert` و `remove` و `sort`. يجب أن تُرجع طريقة `sort` مصفوفة بكل العناصر في الكومة الصغرى مرتبة من الأصغر إلى الأكبر.

# --hints--

يجب أن يوجد هيكل بيانات `MinHeap`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    }
    return typeof test == 'object';
  })()
);
```

يجب أن يحتوي `MinHeap` على طريقة تسمى `insert`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

يجب أن يحتوي `MinHeap` على طريقة تسمى `remove`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

يجب أن يحتوي `MinHeap` على طريقة تسمى `sort`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.sort == 'function';
  })()
);
```

يجب أن تُرجع طريقة `sort` مصفوفة تحتوي على كل العناصر التي أُضيفت إلى الكومة الصغرى بترتيب مرتب.

```js
assert(
  (() => {
    if (typeof MinHeap === 'undefined') {
      return false;
    }

    const heap = new MinHeap();
    const arr = createRandomArray(25);

    for (let i of arr) {
      heap.insert(i);
    }

    const result = heap.sort();
    arr.sort((a, b) => a - b);

    for (let i = 0; i < arr.length; i++) {
      if (arr[i] !== result[i]) {
        return false;
      }
    }
    return true;
  })()
);
```

# --seed--

## --seed-contents--

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
// Generate a randomly filled array
function createRandomArray(size = 5){
  let a = new Array(size);
  for(let i = 0; i < size; i++)
    a[i] = Math.floor(Math.random() * 100);
  
  return a;
}
const array = createRandomArray(25);

var MinHeap = function() {
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
