---
id: 587d8255367417b2b2512c73
title: استخدام معامل النشر والملاحظات لتكامل ES5 Set()
challengeType: 1
forumTopicId: 301720
dashedName: use-spread-and-notes-for-es5-set-integration
---

# --description--

هل تتذكر معامل النشر `...` في ES6؟

يمكن لـ `...` أخذ الكائنات القابلة للتكرار في ES6 وتحويلها إلى مصفوفات.

لننشئ مجموعة Set، ولنستعرض دالة النشر.

```js
var set = new Set([1,2,3]);
var setToArr = [...set]
console.log(setToArr) // returns [ 1, 2, 3 ]
```

# --instructions--

في هذا التمرين، سنمرر كائن Set إلى الدالة `checkSet`. يجب أن تُرجع مصفوفة تحتوي على قيم المجموعة.

الآن تعلمت بنجاح كيفية استخدام كائن ES6 `Set()`، عمل رائع!

# --hints--

يجب أن تُرجع `checkSet(new Set([1,2,3,4,5,6,7])` المصفوفة `[1, 2, 3, 4, 5, 6, 7]`.

```js
var test = checkSet(new Set([1, 2, 3, 4, 5, 6, 7]));
assert.deepEqual(test, [1, 2, 3, 4, 5, 6, 7]);
```

# --seed--

## --seed-contents--

```js
function checkSet(set){
   // Only change code below this line

   // Only change code above this line
}
```

# --solutions--

```js
function checkSet(set){
return [...set];}
```
