---
id: 587d8257367417b2b2512c7b
title: إضافة عنصر جديد إلى شجرة البحث الثنائي
challengeType: 1
forumTopicId: 301618
dashedName: add-a-new-element-to-a-binary-search-tree
---

# --description--

سلسلة التحديات هذه ستُعرّفك على هيكل البيانات الشجري. الأشجار هي هيكل بيانات مهم ومتعدد الاستخدامات في علوم الحاسوب. بالطبع، اسمها يأتي من كونها عند تصورها تبدو مشابهة جدًا للأشجار التي نعرفها في الطبيعة. يبدأ هيكل البيانات الشجري بعقدة واحدة، تُسمى عادةً الجذر، ومن هنا تتفرع إلى عقد إضافية، قد يكون لكل منها عقد فرعية أخرى، وهكذا. عادةً ما يُصوّر هيكل البيانات مع عقدة الجذر في الأعلى؛ يمكنك التفكير فيه كشجرة طبيعية مقلوبة رأسًا على عقب.

أولًا، دعنا نصف بعض المصطلحات الشائعة التي سنواجهها مع الأشجار. عقدة الجذر هي أعلى الشجرة. نقاط البيانات في الشجرة تُسمى عقدًا. العقد التي لها فروع تؤدي إلى عقد أخرى تُسمى عقدة الأصل للعقدة التي يؤدي إليها الفرع (العقدة الفرعية). هناك مصطلحات عائلية أكثر تعقيدًا كما قد تتوقع. الشجرة الفرعية تشير إلى جميع العناصر التابعة لعقدة معينة، والفروع قد تُسمى حواف، والعقد الورقية هي العقد في نهاية الشجرة التي لا تحتوي على أطفال. وأخيرًا، لاحظ أن الأشجار هي هياكل بيانات متكررة بطبيعتها. أي أن أي أطفال لعقدة هم آباء لشجرتهم الفرعية الخاصة، وهكذا. الطبيعة المتكررة للأشجار مهمة لفهمها عند تصميم الخوارزميات للعمليات الشجرية الشائعة.

لنبدأ، سنناقش نوعًا معينًا من الأشجار، الشجرة الثنائية. في الواقع، سنناقش نوعًا معينًا من الشجرة الثنائية، وهي شجرة البحث الثنائي. دعنا نصف ما يعنيه هذا. بينما يمكن لهيكل البيانات الشجري أن يحتوي على أي عدد من الفروع في عقدة واحدة، الشجرة الثنائية يمكن أن تحتوي فقط على فرعين لكل عقدة. علاوة على ذلك، شجرة البحث الثنائي مرتبة بالنسبة للشجرات الفرعية، بحيث تكون قيمة كل عقدة في الشجرة الفرعية اليسرى أقل من أو تساوي قيمة عقدة الأصل، وقيمة كل عقدة في الشجرة الفرعية اليمنى أكبر من أو تساوي قيمة عقدة الأصل. من المفيد جدًا تصور هذه العلاقة لفهمها بشكل أفضل:

<div style='width: 100%; display: flex; justify-content: center; align-items: center;'><img alt="an example of a binary search tree" style='width: 100%; max-width: 350px; background-color: var(--gray-05);' src='https://user-images.githubusercontent.com/18563015/32136009-1e665d98-bbd6-11e7-9133-63184f9f9182.png'></div>

الآن، هذه العلاقة المرتبة سهلة جدًا للرؤية. لاحظ أن كل قيمة على يسار 8، عقدة الجذر، أقل من 8، وكل قيمة على اليمين أكبر من 8. ولاحظ أيضًا أن هذه العلاقة تنطبق على كل الشجرات الفرعية كذلك. على سبيل المثال، أول عقدة فرعية على اليسار هي شجرة فرعية. 3 هي عقدة الأصل، ولها عقدتان فرعيتان بالضبط — وفقًا لقواعد شجرة البحث الثنائي، نعلم دون حتى النظر أن العقدة الفرعية اليسرى لهذه العقدة (وأي من أطفالها) ستكون أقل من 3، والعقدة الفرعية اليمنى (وأي من أطفالها) ستكون أكبر من 3 (ولكن أيضًا أقل من قيمة جذر الهيكل)، وهكذا.

شجرات البحث الثنائي هي هياكل بيانات شائعة ومفيدة جدًا لأنها توفر زمنًا لوغاريتميًا في الحالة المتوسطة لعدة عمليات شائعة مثل البحث، والإدخال، والحذف.

# --instructions--

سنبدأ ببساطة. لقد عرّفنا هيكلًا أساسيًا لشجرة البحث الثنائي هنا بالإضافة إلى دالة لإنشاء العقد لشجرتنا. لاحظ أن كل عقدة قد تحتوي على قيمة يسرى وقيمة يمنى. سيتم تعيين شجرات فرعية كعقد فرعية إذا وجدت. في شجرة البحث الثنائي، ستنشئ طريقة لإضافة قيم جديدة إلى الشجرة. يجب أن تُسمى الطريقة `add` ويجب أن تقبل قيمة صحيحة لإضافتها إلى الشجرة. احرص على الحفاظ على الثابت في شجرة البحث الثنائي: يجب أن تكون القيمة في كل عقدة فرعية يسرى أقل من أو تساوي قيمة عقدة الأصل، والقيمة في كل عقدة فرعية يمنى أكبر من أو تساوي قيمة عقدة الأصل. هنا، دعنا نجعل الشجرة لا تقبل القيم المكررة. إذا حاولنا إضافة قيمة موجودة بالفعل، يجب أن تُرجع الطريقة `null`. وإلا، إذا نجح الإضافة، يجب أن تُرجع `undefined`.

**تلميح:** الأشجار هي هياكل بيانات متكررة بطبيعتها!

# --hints--

يجب أن يوجد هيكل بيانات `BinarySearchTree`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

يجب أن تحتوي شجرة البحث الثنائي على طريقة تُسمى `add`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

يجب أن تضيف طريقة الإضافة العناصر وفقًا لقواعد شجرة البحث الثنائي.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.add !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    const expectedResult = [1, 4, 7, 8, 34, 45, 73, 87];
    const result = test.inOrder();
    return expectedResult.toString() === result.toString();
  })()
);
```

إضافة عنصر موجود بالفعل يجب أن تُرجع `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.add !== 'function') {
      return false;
    }
    test.add(4);
    return test.add(4) == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    inOrder() {
      if (!this.root) {
        return null;
      }
      var result = new Array();
      function traverseInOrder(node) {
        node.left && traverseInOrder(node.left);
        result.push(node.value);
        node.right && traverseInOrder(node.right);
      }
      traverseInOrder(this.root);
      return result;
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.add = function(element) {
    let current = this.root;
    if (!current) {
      this.root = new Node(element);
      return;
    } else {
      const searchTree = function(current) {
        if (current.value > element) {
          if (current.left) {
            return searchTree(current.left);
          } else {
            current.left = new Node(element);
            return;
          }
        } else if (current.value < element) {
          if (current.right) {
            return searchTree(current.right);
          } else {
            current.right = new Node(element);
            return;
          }
        } else {
          return null;
        }
      };
      return searchTree(current);
    }
  };
}
```
