---
id: 587d8259367417b2b2512c84
title: إنشاء شجرة بحث شجرة Trie
challengeType: 1
forumTopicId: 301634
dashedName: create-a-trie-search-tree
---

# --description--

سننتقل هنا من أشجار البحث الثنائية إلى نوع آخر من هياكل الأشجار يُسمى شجرة Trie. شجرة Trie هي شجرة بحث مرتبة تُستخدم عادةً لحفظ السلاسل النصية، أو بشكل أكثر عمومية المصفوفات الترابطية أو مجموعات البيانات الديناميكية التي تكون مفاتيحها سلاسل نصية. هي ممتازة في تخزين مجموعات البيانات عندما تتداخل بادئات العديد من المفاتيح، مثل جميع الكلمات في القاموس. على عكس الشجرة الثنائية، لا ترتبط العقد بقيم فعلية. بدلاً من ذلك، يمثل المسار إلى العقدة مفتاحًا محددًا. على سبيل المثال، إذا أردنا تخزين السلسلة النصية code في شجرة Trie، سيكون لدينا أربع عقد، واحدة لكل حرف: c — o — d — e. اتباع هذا المسار عبر كل هذه العقد سينتج عنه السلسلة النصية code — هذا المسار هو المفتاح الذي خزناه. ثم، إذا أردنا إضافة السلسلة النصية coding، ستشارك أول ثلاث عقد من كلمة code قبل أن تتفرع بعد الحرف d. بهذه الطريقة، يمكن تخزين مجموعات البيانات الكبيرة بشكل مضغوط جدًا. بالإضافة إلى ذلك، يمكن أن يكون البحث سريعًا جدًا لأنه محدود فعليًا بطول السلسلة النصية التي تخزنها. علاوة على ذلك، على عكس الأشجار الثنائية، يمكن للعقدة أن تخزن أي عدد من العقد الفرعية. كما قد تكون خمنت من المثال أعلاه، تُخزن بيانات وصفية عادةً في العقد التي تمثل نهاية مفتاح حتى يمكن استرجاع ذلك المفتاح في عمليات التجوال اللاحقة. على سبيل المثال، إذا أضفنا كلمة codes في مثالنا أعلاه، سنحتاج إلى طريقة لمعرفة أن الحرف e في كلمة code يمثل نهاية مفتاح تم إدخاله سابقًا. وإلا، ستُفقد هذه المعلومة فعليًا عند إضافة كلمة codes.

# --instructions--

لننشئ شجرة Trie لتخزين الكلمات. ستقبل الكلمات عبر طريقة `add` وتخزنها في هيكل بيانات شجرة Trie. كما ستسمح لنا بالاستعلام عما إذا كانت سلسلة نصية معينة كلمة باستخدام طريقة `isWord`، واسترجاع جميع الكلمات المدخلة في الشجرة باستخدام طريقة `print`. يجب أن تُرجع `isWord` قيمة منطقية، ويجب أن تُرجع `print` مصفوفة بكل هذه الكلمات كسلاسل نصية. للتحقق من تنفيذ هذا الهيكل بشكل صحيح، وفرنا هيكل `Node` لكل عقدة في الشجرة. ستكون كل عقدة كائنًا يحتوي على خاصية `keys` وهي كائن Map في JavaScript. هذا سيحتوي على الحروف الفردية التي تمثل مفاتيح صالحة لكل عقدة. كما أنشأنا خاصية `end` على العقد يمكن تعيينها إلى `true` إذا كانت العقدة تمثل نهاية كلمة.

# --hints--

يجب أن تحتوي شجرة `Trie` على طريقة `add`.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

يجب أن تحتوي شجرة `Trie` على طريقة `print`.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

يجب أن تحتوي شجرة `Trie` على طريقة `isWord`.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    return typeof test.isWord == 'function';
  })()
);
```

يجب أن تُرجع طريقة `print` جميع العناصر المضافة إلى شجرة Trie كسلاسل نصية في مصفوفة.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    test.add('jump');
    test.add('jumps');
    test.add('jumped');
    test.add('house');
    test.add('mouse');
    var added = test.print();
    return (
      added.indexOf('jump') != -1 &&
      added.indexOf('jumps') != -1 &&
      added.indexOf('jumped') != -1 &&
      added.indexOf('house') != -1 &&
      added.indexOf('mouse') != -1 &&
      added.length == 5
    );
  })()
);
```

يجب أن تُرجع طريقة `isWord` القيمة `true` فقط للكلمات المضافة إلى شجرة Trie و`false` لجميع الكلمات الأخرى.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    test.add('hop');
    test.add('hops');
    test.add('hopped');
    test.add('hoppy');
    test.add('hope');
    return (
      test.isWord('hop') &&
      !test.isWord('ho') &&
      test.isWord('hopped') &&
      !test.isWord('hopp') &&
      test.isWord('hoppy') &&
      !test.isWord('hoping')
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
var Node = function() {
  this.keys = new Map();
  this.end = false;
  this.setEnd = function() {
    this.end = true;
  };
  this.isEnd = function() {
    return this.end;
  };
};
var Trie = function() {
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
