---
id: 587d8251367417b2b2512c64
title: البحث داخل قائمة مرتبطة
challengeType: 1
forumTopicId: 301715
dashedName: search-within-a-linked-list
---

# --description--

لنُضف بعض الطرق المفيدة إلى صنف القائمة المرتبطة. ألن يكون من المفيد أن نعرف إذا كانت قائمتنا فارغة أم لا، كما في أصناف `Stack` و `Queue`؟

يجب أن نتمكن أيضًا من إيجاد عناصر محددة في القائمة المرتبطة. التنقل عبر هياكل البيانات هو أمر ستحتاج إلى الكثير من التدريب عليه! لننشئ طريقة `indexOf` تأخذ `element` كمعلمة، وتُرجع `index` ذلك العنصر في القائمة المرتبطة. إذا لم يُعثر على العنصر في القائمة المرتبطة، أرجع `-1`.

لننفذ أيضًا طريقة تقوم بالعكس: طريقة `elementAt` تأخذ `index` كمعلمة وتُرجع `element` عند هذا الـ `index`. إذا لم يُعثر على `element`، أرجع `undefined`.

# --instructions--

اكتب طريقة `isEmpty` تتحقق مما إذا كانت القائمة المرتبطة فارغة، وطريقة `indexOf` تُرجع فهرس `index` عنصر معين، وطريقة `elementAt` تُرجع العنصر `element` عند فهرس `index` معين.

# --hints--

يجب أن يحتوي صنف `LinkedList` على طريقة `isEmpty`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.isEmpty === 'function';
  })()
);
```

يجب أن تُرجع طريقة `isEmpty` القيمة `false` عندما يكون هناك عنصر واحد على الأقل في القائمة المرتبطة.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.isEmpty() === false;
  })()
);
```

يجب أن تُرجع طريقة `isEmpty` القيمة `true` عندما لا توجد عناصر في القائمة المرتبطة.

```js
assert(
  (function () {
    var test = new LinkedList();
    return test.isEmpty() === true;
  })()
);
```

يجب أن يحتوي صنف `LinkedList` على طريقة `indexOf`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.indexOf === 'function';
  })()
);
```

يجب أن تُرجع طريقة `indexOf` فهرس عنصر معين موجود في القائمة المرتبطة.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.indexOf('cat') === 0;
  })()
);
```

يجب أن تُرجع طريقة `indexOf` القيمة `-1` إذا لم يُعثر على العنصر المعطى في القائمة المرتبطة.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.indexOf('pony') === -1;
  })()
);
```

يجب أن يحتوي صنف `LinkedList` على طريقة `elementAt`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.elementAt === 'function';
  })()
);
```

يجب أن تُرجع طريقة `elementAt` العنصر الموجود عند فهرس معين في القائمة المرتبطة.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.elementAt(1) === 'dog';
  })()
);
```

يجب أن تُرجع طريقة `elementAt` القيمة `undefined` إذا لم يُعثر على العنصر عند الفهرس المعطى في القائمة المرتبطة.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.elementAt(5) === undefined;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    var currentNode = head;
    var previousNode;
    if(currentNode.element === element){
      head = currentNode.next;
    } else {
      while(currentNode.element !== element) {
        previousNode = currentNode;
        currentNode = currentNode.next;
      }

      previousNode.next = currentNode.next;
    }

    length --;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
        var currentNode = head;

        while(currentNode.next){
            currentNode  = currentNode.next;
        }

        currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    var currentNode = head;
    var previousNode;
    if(currentNode.element === element){
        head = currentNode.next;
    } else {
        while(currentNode.element !== element) {
            previousNode = currentNode;
            currentNode = currentNode.next;
        }

        previousNode.next = currentNode.next;
    }

    length --;
  };

  this.indexOf = function(element) {
    if (head === null) return -1

    let current = head;
    let index = 0;

    while (current.element !== element && current.next !== null) {
      current = current.next;
      index++
    }

    if (current.element !== element && current.next === null) {
      return -1
    }

    return index;
  }

  this.elementAt = function(index) {
    if (head === null) return undefined;

    let current = head;
    let currentIndex = 0;

    while (currentIndex !== index && current.next !== null) {
      current = current.next;
      currentIndex++
    }

    if (currentIndex !== index && current.next === null) {
      return undefined;
    }

    return current.element;
  }

  this.isEmpty = function() {
    return length === 0;
  }
}
```
