---
id: 587d825b367417b2b2512c8b
title: إزالة عنصر من حزمة قصوى
challengeType: 1
forumTopicId: 301710
dashedName: remove-an-element-from-a-max-heap
---

# --description--

الآن بعد أن أصبح بإمكاننا إضافة عناصر إلى الحزمة، لنرَ كيف يمكننا إزالة العناصر. تتطلب عمليتا الإزالة والإدخال منطقًا مشابهًا. في الحزمة القصوى، عادةً ما تريد إزالة أكبر قيمة، وهذا يعني ببساطة استخراجها من جذر الشجرة. هذا سيكسر خاصية الحزمة في شجرتنا، لذا يجب إعادة تأسيسها بطريقة ما. عادةً، يتم ذلك في الحزمة القصوى بالطريقة التالية:

<ol>
<li>نقل العنصر الأخير في الحزمة إلى موقع الجذر.</li>
  <li>إذا كان أي من أبناء الجذر أكبر منه، فبدّل الجذر مع الابن ذي القيمة الأكبر.</li>
  <li>استمر في التبديل حتى يصبح الأصل أكبر من كلا الابنين أو تصل إلى المستوى الأخير في الشجرة.</li>
</ol>

# --instructions--

التعليمات: أضف طريقة إلى الحزمة القصوى لدينا تُسمى `remove`. يجب أن تُرجع هذه الطريقة أكبر قيمة أُضيفت إلى الحزمة القصوى وتزيلها منها. كما يجب أن تعيد ترتيب الحزمة بحيث تُحافظ على خاصية الحزمة. بعد إزالة عنصر، يجب أن يصبح العنصر الأكبر التالي المتبقي في الحزمة هو الجذر.

# --hints--

يجب أن يوجد هيكل البيانات `MaxHeap`.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    }
    return typeof test == 'object';
  })()
);
```

يجب أن يحتوي `MaxHeap` على طريقة تُسمى `print`.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

يجب أن يحتوي `MaxHeap` على طريقة تُسمى `insert`.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

يجب أن يحتوي `MaxHeap` على طريقة تُسمى `remove`.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

يجب أن تزيل طريقة `remove` أكبر عنصر من الحزمة القصوى مع الحفاظ على خاصية الحزمة القصوى.

```js
function isHeap(arr, i, n) {
    if( arr[i] < arr[2 * i + 1] || arr[i] < arr[2 * i + 2] ){
        return false;
    }
    if (i > (n - 1) / 2) {
        return true;
    }
    if (isHeap(arr, 2 * i + 1, n) && isHeap(arr, 2 * i + 2, n)) {
        return true;
    }
    return false;
}


assert(
  (function () {
    let test = false;

    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    
    let max = Infinity;
    const [result, vals] = [[], [9, 3, 5, 2, 15, 3, 7, 12, 7, 10, 90]];

    vals.forEach((val) => test.insert(val));

    for (let i = 0; i < vals.length; i++) {
      const curHeap = test.print();
      const arr = curHeap[0] === null ? curHeap.slice(1) : curHeap;
      
      if (!isHeap(arr, 0, arr.length - 1)) {
        return false;
      }

      const removed = test.remove();
      if (!vals.includes(removed)) return false;
      if (removed > max) return false
      max = removed;
      result.push(removed);
    }
    
    for (let i = 0; i < vals.length; i++) {
      if (!result.includes(vals[i])) {
        return false;
      }
    }
    
    return true;
  })()
);
```

# --seed--

## --seed-contents--

```js
const MaxHeap = function () {
  this.heap = [];
  this.parent = index => {
    return Math.floor((index - 1) / 2);
  }
  this.insert = element => {
    this.heap.push(element);
    this.heapifyUp(this.heap.length - 1);
  }
  this.heapifyUp = index => {
    let currentIndex = index,
    parentIndex = this.parent(currentIndex);
    while (currentIndex > 0 && this.heap[currentIndex] > this.heap[parentIndex]) {
      this.swap(currentIndex, parentIndex);
      currentIndex = parentIndex;
      parentIndex = this.parent(parentIndex);
    }
  }
  this.swap = (index1, index2) => {
    [this.heap[index1], this.heap[index2]] = [this.heap[index2], this.heap[index1]];
  }
  this.print = () => {
    return this.heap;
  }
  // Only change code below this line

  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
