---
id: 587d8251367417b2b2512c63
title: إزالة عناصر من قائمة مرتبطة
challengeType: 1
forumTopicId: 301712
dashedName: remove-elements-from-a-linked-list
---

# --description--

الطريقة المهمة التالية التي يحتاجها أي تنفيذ لقائمة مرتبطة هي طريقة `remove`. يجب أن تأخذ هذه الطريقة العنصر الذي نريد إزالته كمعلمة، ثم تبحث في القائمة لإيجاد العقدة التي تحتوي على هذا العنصر وإزالتها.

عندما نزيل عقدة من قائمة مرتبطة، من المهم ألا نتسبب عن غير قصد في ترك بقية القائمة بدون اتصال. تذكر أن خاصية `next` لكل عقدة تشير إلى العقدة التي تليها في القائمة. إذا كنا نزيل العنصر الأوسط، على سبيل المثال، يجب أن نتأكد من وجود اتصال من خاصية `next` للعقدة السابقة لذلك العنصر إلى خاصية `next` للعنصر الأوسط (وهي العقدة التالية في القائمة!).

قد يبدو هذا مربكًا جدًا، لذا لنعد إلى مثال صف الكونغا لنحصل على نموذج مفاهيمي جيد. تخيل نفسك في صف الكونغا، والشخص الذي أمامك مباشرة يغادر الصف. الشخص الذي غادر لم يعد يمسك بأي شخص في الصف، وأنت لم تعد تمسك بالشخص الذي غادر. تتقدم خطوة وتضع يديك على الشخص التالي الذي تراه.

إذا كان العنصر الذي نريد إزالته هو العنصر `head`، نُعيد تعيين `head` إلى العقدة الثانية في القائمة المرتبطة.

# --instructions--

اكتب طريقة `remove` تأخذ عنصرًا وتزيله من القائمة المرتبطة.

**ملاحظة:** يجب أن يقل طول `length` القائمة بمقدار واحد في كل مرة يُزال فيها عنصر من القائمة المرتبطة.

# --hints--

يجب أن يحتوي صنف `LinkedList` على طريقة `remove`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.remove === 'function';
  })()
);
```

يجب أن تعيد طريقة `remove` تعيين `head` إلى العقدة الثانية عند إزالة العقدة الأولى.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.remove('cat');
    return test.head().element === 'dog';
  })()
);
```

يجب أن تقلل طريقة `remove` طول `length` القائمة المرتبطة بمقدار واحد لكل عقدة تُزال.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('hamster');
    test.remove('cat');
    test.remove('fish');
    return test.size() === 2;
  })()
);
```

يجب أن تعيد طريقة `remove` تعيين المرجع الخاص بالعقدة السابقة للعقدة التي أُزيلت إلى المرجع `next` الخاص بالعقدة التي أُزيلت.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('snake');
    test.add('kitten');
    test.remove('snake');
    return test.head().next.next.element === 'kitten';
  })()
);
```

يجب ألا تغير طريقة `remove` القائمة المرتبطة إذا لم يكن العنصر موجودًا في القائمة.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.remove('elephant');
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
        var currentNode = head;

        while(currentNode.next){
            currentNode  = currentNode.next;
        }

        currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    if (head === null) {
      return;
    }
    var previous;
    var currentNode = head;

    while (currentNode.next !== null && currentNode.element !== element) {
      previous = currentNode;
      currentNode = currentNode.next;
    }
    
    if (currentNode.next === null && currentNode.element !== element) {
      return;
    }
    else if (previous) {
      previous.next = currentNode.next;
    } else {
      head = currentNode.next;
    }

    length--;
  };
} 
```
