---
id: 587d8252367417b2b2512c67
title: إضافة عناصر عند فهرس محدد في قائمة مرتبطة
challengeType: 1
forumTopicId: 301619
dashedName: add-elements-at-a-specific-index-in-a-linked-list
---

# --description--

لننشئ طريقة addAt(index,element) تضيف عنصرًا عند فهرس محدد. تمامًا كما نزيل العناصر عند فهرس معين، نحتاج إلى تتبع currentIndex أثناء التنقل في القائمة المرتبطة. عندما يتطابق currentIndex مع الفهرس المعطى، يجب إعادة تعيين خاصية next للعقدة السابقة لتشير إلى العقدة الجديدة المضافة. ويجب أن تشير العقدة الجديدة إلى العقدة التالية في currentIndex. بالعودة إلى مثال صف الكونغا، هناك شخص جديد يريد الانضمام إلى الصف، لكنه يريد الانضمام في المنتصف. أنت في منتصف الصف، فتترك يد الشخص الذي أمامك. يأتي الشخص الجديد ويضع يديه على الشخص الذي كنت تمسك يديه، وأنت الآن تمسك يدي الشخص الجديد.

# --instructions--

أنشئ طريقة `addAt(index,element)` تضيف عنصرًا عند فهرس معين. أرجع false إذا لم يمكن إضافة العنصر. **ملاحظة:** تذكر التحقق مما إذا كان الفهرس المعطى سالبًا أو أطول من طول القائمة المرتبطة.

# --hints--

يجب أن تعيد طريقة `addAt` تعيين `head` إلى العقدة الجديدة عندما يكون الفهرس المعطى 0.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.addAt(0, 'fish');
    return test.head().element === 'fish' && test.head().next.element === 'cat';
  })()
);
```

يجب أن تزيد طريقة `addAt` طول القائمة المرتبطة بمقدار واحد لكل عقدة جديدة تضاف إلى القائمة.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.addAt(0, 'cat');
    return test.size() === 3;
  })()
);
```

يجب أن تعيد طريقة `addAt` القيمة `false` إذا لم يمكن إضافة عقدة.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    return test.addAt(4, 'cat') === false;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element) {
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function() {
    return head;
  };

  this.add = function(element) {
    var node = new Node(element);
    if (head === null) {
      head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }
    length++;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if (head === null){
        head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }
    length++;
  };
  this.addAt = function (index, element) {
    if (index > length || index < 0) {
      return false;
    }
    var newNode = new Node(element);
    var currentNode = head;
    if (index === 0) {
      head = newNode;
    } else {
      var previousNode = null;
      var i = 0;
      while (currentNode && i < index) {
        previousNode = currentNode;
        currentNode = currentNode.next;
        i++;
      }
      previousNode.next = newNode;
    }
    newNode.next = currentNode;
    length++;
  }
}
```
