---
id: 587d8257367417b2b2512c7c
title: التحقق من وجود عنصر في شجرة البحث الثنائي
challengeType: 1
forumTopicId: 301623
dashedName: check-if-an-element-is-present-in-a-binary-search-tree
---

# --description--

الآن بعد أن أصبح لدينا فكرة عامة عن شجرة البحث الثنائي، دعنا نتعمق فيها أكثر قليلًا. توفر أشجار البحث الثنائي زمنًا لوغاريتميًا للعمليات الشائعة مثل البحث، والإدراج، والحذف في الحالة المتوسطة، وزمنًا خطيًا في أسوأ الحالات. لماذا ذلك؟ كل واحدة من هذه العمليات الأساسية تتطلب منا إيجاد عنصر في الشجرة (أو في حالة الإدراج إيجاد المكان المناسب له)، وبسبب بنية الشجرة، عند كل عقدة أصل نتفرع يسارًا أو يمينًا، ونستبعد فعليًا نصف حجم الشجرة المتبقية. هذا يجعل البحث يتناسب مع لوغاريتم عدد العقد في الشجرة، مما يؤدي إلى زمن لوغاريتمي لهذه العمليات في الحالة المتوسطة. حسنًا، ماذا عن أسوأ الحالات؟ تخيل بناء شجرة من القيم التالية، بإضافتها من اليسار إلى اليمين: `10`، `12`، `17`، `25`. باتباع قواعد شجرة البحث الثنائي، سنضيف `12` إلى يمين `10`، و`17` إلى يمينها، و`25` إلى يمينها أيضًا. الآن أصبحت شجرتنا تشبه قائمة مرتبطة، والبحث عن `25` يتطلب المرور على جميع العناصر بشكل خطي. لهذا السبب يكون الزمن خطيًا في أسوأ الحالات. المشكلة هنا أن الشجرة غير متوازنة. سنتعرف أكثر على معنى ذلك في التحديات التالية.

# --instructions--

في هذا التحدٍّ، سننشئ أداة مساعدة لشجرتنا. اكتب طريقة `isPresent` التي تأخذ قيمة صحيحة كمُدخَل وتُرجع قيمة منطقية تدل على وجود أو عدم وجود تلك القيمة في شجرة البحث الثنائي.

# --hints--

يجب أن يوجد هيكل البيانات `BinarySearchTree`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

يجب أن تحتوي شجرة البحث الثنائي على طريقة تسمى `isPresent`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.isPresent == 'function';
  })()
);
```

يجب أن تتحقق طريقة `isPresent` بشكل صحيح من وجود أو عدم وجود العناصر المضافة إلى الشجرة.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isPresent !== 'function') {
      return false;
    }
    test.add(4);
    test.add(7);
    test.add(411);
    test.add(452);
    return (
      test.isPresent(452) &&
      test.isPresent(411) &&
      test.isPresent(7) &&
      !test.isPresent(100)
    );
  })()
);
```

يجب أن تتعامل `isPresent` مع الحالات التي تكون فيها الشجرة فارغة.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isPresent !== 'function') {
      return false;
    }
    return test.isPresent(5) == false;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = (tree) => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.isPresent = function (value) {
    var current = this.root
    while (current) {
      if (value === current.value) {
        return true;
      }
      current = value < current.value ? current.left : current.right;
    }
    return false;
  }
}
```
