---
id: 587d825c367417b2b2512c90
title: البحث بالاتساع أولاً
challengeType: 1
forumTopicId: 301622
dashedName: breadth-first-search
---

# --description--

حتى الآن، تعلمنا طرقًا مختلفة لإنشاء تمثيلات للرسوم البيانية. ماذا بعد؟ سؤال طبيعي هو ما هي المسافات بين أي عقدتين في الرسم البياني؟ هنا تدخل <dfn>خوارزميات اجتياز الرسم البياني</dfn>.

<dfn>خوارزميات الاجتياز</dfn> هي خوارزميات لتجوال أو زيارة العقد في الرسم البياني. أحد أنواع خوارزميات الاجتياز هو خوارزمية البحث بالاتساع أولاً.

تبدأ هذه الخوارزمية من عقدة واحدة وتزور جميع جيرانها الذين يبعدون حافة واحدة. ثم تتابع لزيارة جيران كل منهم وهكذا حتى يتم الوصول إلى جميع العقد.

هيكل بيانات مهم يساعد في تنفيذ خوارزمية البحث بالاتساع أولاً هو الطابور. هو مصفوفة يمكنك إضافة عناصر إلى أحد طرفيها وإزالة عناصر من الطرف الآخر. يُعرف هذا أيضًا باسم هيكل بيانات <dfn>FIFO</dfn> أو <dfn>الأول داخل الأول خارج</dfn>.

بصريًا، هذا ما تفعله الخوارزمية:
<img alt="animation showing the breadth first search algorithm" src='https://cdn.freecodecamp.org/curriculum/coding-interview-prep/breadth-first-search.gif'>

التظليل الرمادي يمثل إضافة عقدة إلى الطابور والتظليل الأسود يمثل إزالة عقدة من الطابور. لاحظ كيف في كل مرة تُزال عقدة من الطابور (تصبح العقدة سوداء)، تُضاف جميع جيرانها إلى الطابور (تصبح العقدة رمادية).

لتنفيذ هذه الخوارزمية، ستحتاج إلى إدخال هيكل الرسم البياني وعقدة تريد البدء منها.

أولًا، يجب أن تكون على علم بالمسافات من العقدة البداية، أو عدد الحواف التي تبعدها. ابدأ جميع المسافات برقم كبير مثل `Infinity`. هذا يمنع مشاكل العد عندما لا يمكن الوصول إلى عقدة من العقدة البداية. بعد ذلك، انتقل من العقدة البداية إلى جيرانها. هؤلاء الجيران يبعدون حافة واحدة، وعند هذه النقطة يجب إضافة وحدة واحدة إلى المسافات التي تتابعها.

# --instructions--

اكتب دالة `bfs()` تأخذ رسمًا بيانيًا على شكل مصفوفة مجاورة (مصفوفة ثنائية الأبعاد) ومسمى عقدة الجذر كمعلمات. سيكون مسمى العقدة هو القيمة الصحيحة للعقدة بين `0` و `n - 1`، حيث `n` هو العدد الإجمالي للعقد في الرسم البياني.

ستُخرج دالتك كائن JavaScript يحتوي على أزواج المفتاح-القيمة مع العقدة والمسافة من الجذر. إذا لم يمكن الوصول إلى العقدة، يجب أن تكون المسافة `Infinity`.

# --hints--

الرسم البياني المدخل `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` مع عقدة بداية `1` يجب أن يُرجع `{0: 1, 1: 0, 2: 1, 3: 2}`

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
var results = bfs(graph, 1);
assert.deepEqual(results, { 0: 1, 1: 0, 2: 1, 3: 2 });
```

الرسم البياني المدخل `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` مع عقدة بداية `1` يجب أن يُرجع `{0: 1, 1: 0, 2: 1, 3: Infinity}`

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 0],
  [0, 0, 0, 0]
];
var results = bfs(graph, 1);
 assert.deepEqual(results, { 0: 1, 1: 0, 2: 1, 3: Infinity });
```

الرسم البياني المدخل `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` مع عقدة بداية `0` يجب أن يُرجع `{0: 0, 1: 1, 2: 2, 3: 3}`

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
var results = bfs(graph, 0);
 assert.deepEqual(results, { 0: 0, 1: 1, 2: 2, 3: 3 });
```

الرسم البياني المدخل `[[0, 1], [1, 0]]` مع عقدة بداية `0` يجب أن يُرجع `{0: 0, 1: 1}`

```js
var graph = [
  [0, 1],
  [1, 0]
];
var results = bfs(graph, 0);
assert.deepEqual(results, { 0: 0, 1: 1 });
```

# --seed--

## --seed-contents--

```js
function bfs(graph, root) {
  var nodesLen = {};

  return nodesLen;
};

var exBFSGraph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
console.log(bfs(exBFSGraph, 3));
```

# --solutions--

```js
function bfs(graph, root) {
  var nodesLen = {};
  // Set all distances to infinity
  for (var i = 0; i < graph.length; i++) {
    nodesLen[i] = Infinity;
  }
  nodesLen[root] = 0; // ...except root node
  var queue = [root]; // Keep track of nodes to visit
  var current; // Current node traversing
  // Keep on going until no more nodes to traverse
  while (queue.length !== 0) {
    current = queue.shift();
    // Get adjacent nodes from current node
    var curConnected = graph[current]; // Get layer of edges from current
    var neighborIdx = []; // List of nodes with edges
    var idx = curConnected.indexOf(1); // Get first edge connection
    while (idx !== -1) {
      neighborIdx.push(idx); // Add to list of neighbors
      idx = curConnected.indexOf(1, idx + 1); // Keep on searching
    }
    // Loop through neighbors and get lengths
    for (var j = 0; j < neighborIdx.length; j++) {
      // Increment distance for nodes traversed
      if (nodesLen[neighborIdx[j]] === Infinity) {
        nodesLen[neighborIdx[j]] = nodesLen[current] + 1;
        queue.push(neighborIdx[j]); // Add new neighbors to queue
      }
    }
  }
  return nodesLen;
}
```
