---
id: 587d8257367417b2b2512c7e
title: استخدام البحث بالعمق أولاً في شجرة البحث الثنائي
challengeType: 1
forumTopicId: 301719
dashedName: use-depth-first-search-in-a-binary-search-tree
---

# --description--

نعرف كيف نبحث في شجرة البحث الثنائي عن قيمة محددة. لكن ماذا لو أردنا فقط استكشاف الشجرة بأكملها؟ أو ماذا لو لم تكن الشجرة مرتبة ونحتاج فقط للبحث عن قيمة؟ هنا سنقدم بعض طرق اجتياز الشجرة التي يمكن استخدامها لاستكشاف هياكل البيانات الشجرية. أولها هو البحث بالعمق أولاً. في البحث بالعمق أولاً، تُستكشف شجرة فرعية معينة بأقصى عمق ممكن قبل أن يستمر البحث في شجرة فرعية أخرى. هناك ثلاث طرق يمكن تنفيذ ذلك بها: بالترتيب: ابدأ البحث من العقدة اليسرى الأكثر وانتهِ عند العقدة اليمنى الأكثر. ما قبل الترتيب: استكشف جميع الجذور قبل الأوراق. ما بعد الترتيب: استكشف جميع الأوراق قبل الجذور. كما قد تتوقع، يمكنك اختيار طرق بحث مختلفة حسب نوع البيانات التي تخزنها الشجرة وما تبحث عنه. بالنسبة لشجرة البحث الثنائي، يعيد الاجتياز بالترتيب العقد مرتبة.

# --instructions--

هنا سننشئ هذه الطرق الثلاث للبحث على شجرة البحث الثنائي الخاصة بنا. البحث بالعمق أولاً هو عملية متكررة بطبيعتها تستمر في استكشاف الأشجار الفرعية طالما توجد عقد فرعية. بمجرد فهم هذا المفهوم الأساسي، يمكنك ببساطة إعادة ترتيب ترتيب استكشاف العقد والأشجار الفرعية لإنتاج أي من الطرق الثلاث السابقة. على سبيل المثال، في البحث ما بعد الترتيب نريد أن نستدعي التكرار حتى نصل إلى عقدة ورقية قبل أن نبدأ بإرجاع أي من العقد نفسها، بينما في البحث ما قبل الترتيب نريد إرجاع العقد أولاً، ثم نتابع التكرار نزولاً في الشجرة. عرّف طرق `inorder` و `preorder` و `postorder` على شجرتنا. يجب أن تُرجع كل من هذه الطرق مصفوفة من العناصر التي تمثل اجتياز الشجرة. تأكد من إرجاع القيم الصحيحة في كل عقدة داخل المصفوفة، وليس العقد نفسها. وأخيرًا، أرجع `null` إذا كانت الشجرة فارغة.

# --hints--

يجب أن يوجد هيكل البيانات `BinarySearchTree`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

يجب أن تحتوي شجرة البحث الثنائي على طريقة تسمى `inorder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.inorder == 'function';
  })()
);
```

يجب أن تحتوي شجرة البحث الثنائي على طريقة تسمى `preorder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.preorder == 'function';
  })()
);
```

يجب أن تحتوي شجرة البحث الثنائي على طريقة تسمى `postorder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.postorder == 'function';
  })()
);
```

يجب أن تُرجع طريقة `inorder` مصفوفة من قيم العقد الناتجة عن اجتياز بالترتيب.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.inorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.inorder().join('') == '012345678910';
  })()
);
```

يجب أن تُرجع طريقة `preorder` مصفوفة من قيم العقد الناتجة عن اجتياز ما قبل الترتيب.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.preorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.preorder().join('') == '710325469810';
  })()
);
```

يجب أن تُرجع طريقة `postorder` مصفوفة من قيم العقد الناتجة عن اجتياز ما بعد الترتيب.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.postorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.postorder().join('') == '024653181097';
  })()
);
```

يجب أن تُرجع طريقة `inorder` القيمة `null` للشجرة الفارغة.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.inorder !== 'function') {
      return false;
    }
    return test.inorder() == null;
  })()
);
```

يجب أن تُرجع طريقة `preorder` القيمة `null` للشجرة الفارغة.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.preorder !== 'function') {
      return false;
    }
    return test.preorder() == null;
  })()
);
```

يجب أن تُرجع طريقة `postorder` القيمة `null` للشجرة الفارغة.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.postorder !== 'function') {
      return false;
    }
    return test.postorder() == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      function searchTree(node) {
        if (value < node.value) {
          if (node.left == null) {
            node.left = new Node(value);
            return;
          } else if (node.left != null) {
            return searchTree(node.left);
          }
        } else if (value > node.value) {
          if (node.right == null) {
            node.right = new Node(value);
            return;
          } else if (node.right != null) {
            return searchTree(node.right);
          }
        } else {
          return null;
        }
      }

      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.result = [];

  this.inorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    if (node.left) this.inorder(node.left);
    this.result.push(node.value);
    if (node.right) this.inorder(node.right);
    return this.result;
  };
  this.preorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    this.result.push(node.value);
    if (node.left) this.preorder(node.left);
    if (node.right) this.preorder(node.right);
    return this.result;
  };
  this.postorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    if (node.left) this.postorder(node.left);
    if (node.right) this.postorder(node.right);
    this.result.push(node.value);

    return this.result;
  };
}
```
