---
id: 587d8250367417b2b2512c5f
title: إنشاء صنف مكدس
challengeType: 1
forumTopicId: 301633
dashedName: create-a-stack-class
---

# --description--

في القسم السابق، تحدثنا عن ماهية المكدس وكيف يمكننا استخدام مصفوفة لتمثيل المكدس. في هذا القسم، سننشئ صنف المكدس الخاص بنا. رغم أنه يمكنك استخدام المصفوفات لإنشاء مكدسات، إلا أنه من الأفضل أحيانًا تحديد مقدار التحكم الذي نملكه في مكدساتنا. بالإضافة إلى طريقتي `push` و `pop`، للمكدسات طرق مفيدة أخرى. دعنا نضيف طريقتي `peek` و `isEmpty` و `clear` إلى صنف المكدس.

# --instructions--

اكتب طريقة `push` تضيف عنصرًا إلى أعلى المكدس، وطريقة `pop` تزيل وتُعيد العنصر الموجود في أعلى المكدس، وطريقة `peek` تنظر إلى العنصر الأعلى في المكدس، وطريقة `isEmpty` تتحقق مما إذا كان المكدس فارغًا، وطريقة `clear` تزيل كل العناصر من المكدس. عادةً لا تحتوي المكدسات على هذه الطريقة، لكننا أضفنا طريقة مساعدة `print` تقوم بطباعة المجموعة في وحدة التحكم.

# --hints--

يجب أن يحتوي صنف `Stack` على طريقة `push`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.push === 'function';
  })()
);
```

يجب أن يحتوي صنف `Stack` على طريقة `pop`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.pop === 'function';
  })()
);
```

يجب أن يحتوي صنف `Stack` على طريقة `peek`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.peek === 'function';
  })()
);
```

يجب أن يحتوي صنف `Stack` على طريقة `isEmpty`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.isEmpty === 'function';
  })()
);
```

يجب أن يحتوي صنف `Stack` على طريقة `clear`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.clear === 'function';
  })()
);
```

يجب أن تُعيد طريقة `peek` العنصر الأعلى في المكدس

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    return test.peek() === 'CS50' && test.peek() === 'CS50';
  })()
);
```

يجب أن تزيل طريقة `pop` العنصر الأعلى في المكدس وتُعيده

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    return test.pop() === 'CS50' && test.pop() === 'CS61';
  })()
);
```

يجب أن تُعيد طريقة `isEmpty` القيمة true إذا لم يحتوي المكدس على أي عناصر

```js
assert(
  (function () {
    var test = new Stack();
    return test.isEmpty();
  })()
);
```

يجب أن تزيل طريقة `clear` كل العناصر من المكدس

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    test.clear();
    return test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Stack() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
class Stack {
  constructor() {
    this.collection = [];
  }
  print() {
    console.log(this.collection);
  }
  push(val) {
    this.collection.push(val);
  }
  pop() {
    return this.collection.pop();
  }
  peek() {
    return this.collection[this.collection.length - 1];
  }
  isEmpty() {
    return this.collection.length === 0;
  }
  clear() {
    return (this.collection.length = 0);
  }
}
```
