---
id: 587d8255367417b2b2512c72
title: استخدام .has و .size على مجموعة ES6
challengeType: 1
forumTopicId: 301717
dashedName: use--has-and--size-on-an-es6-set
---

# --description--

لننظر إلى الطريقتين .has و .size المتاحتين على كائن مجموعة ES6.

أولًا، أنشئ مجموعة ES6

```js
var set = new Set([1,2,3]);
```

تتحقق الطريقة .has مما إذا كانت القيمة موجودة داخل المجموعة.

```js
var hasTwo = set.has(2);
```

تعيد الطريقة .size عددًا صحيحًا يمثل حجم المجموعة

```js
var howBig = set.size;
```

# --instructions--

في هذا التمرين، سنمرر مصفوفة وقيمة إلى الدالة checkSet(). يجب أن تنشئ دالتك مجموعة ES6 من المعلمة المصفوفة. تحقق مما إذا كانت المجموعة تحتوي على المعلمة القيمة. احصل على حجم المجموعة. وأعد هاتين القيمتين في مصفوفة.

# --hints--

`checkSet([4, 5, 6], 3)` يجب أن تعيد [ false, 3 ]

```js
var test = checkSet([4, 5, 6], 3);
assert.deepEqual(test, [false, 3]);
```

# --seed--

## --seed-contents--

```js
function checkSet(arrToBeSet, checkValue){

   // Only change code below this line

   // Only change code above this line

}
```

# --solutions--

```js
function checkSet(arrToBeSet, checkValue){
var set = new Set(arrToBeSet);
var result = [
set.has(checkValue),
set.size
];
return result;
}
```
