---
id: 587d825a367417b2b2512c8a
title: إدخال عنصر في حزمة قصوى
challengeType: 1
forumTopicId: 301703
dashedName: insert-an-element-into-a-max-heap
---

# --description--

سننتقل الآن إلى هيكل بيانات شجري آخر، وهو الحزمة الثنائية. الحزمة الثنائية هي شجرة ثنائية مرتبة جزئيًا تحقق خاصية الحزمة. تحدد خاصية الحزمة علاقة بين عقدة الأصل والعقد الفرعية. قد يكون لديك حزمة قصوى، حيث تكون جميع عقد الأصل أكبر من أو تساوي عقدها الفرعية، أو حزمة صغرى حيث العكس صحيح. الحزم الثنائية هي أيضًا أشجار ثنائية كاملة. هذا يعني أن جميع مستويات الشجرة ممتلئة بالكامل وإذا كان المستوى الأخير ممتلئًا جزئيًا فإنه يُملأ من اليسار إلى اليمين.

بينما يمكن تنفيذ الحزم الثنائية كهيكل شجري يحتوي على عقد بها مراجع لليسار واليمين، يسمح الترتيب الجزئي وفقًا لخاصية الحزمة بتمثيل الحزمة باستخدام مصفوفة. العلاقة بين الأصل والأبناء هي ما نهتم به، وبحسابات بسيطة يمكننا حساب الأبناء لأي عقدة أصل والعقدة الأصل لأي عقدة فرعية.

على سبيل المثال، انظر إلى هذا التمثيل المصفوفي لحزمة ثنائية صغرى:

```js
[ 6, 22, 30, 37, 63, 48, 42, 76 ]
```

العقدة الجذرية هي العنصر الأول، `6`. أبناؤها هم `22` و `30`. إذا نظرنا إلى العلاقة بين فهارس المصفوفة لهذه القيم، فإن الأبناء عند الفهرس `i` هم `2 * i + 1` و `2 * i + 2`. وبالمثل، العنصر عند الفهرس `0` هو عقدة الأصل لهذين الابنين عند الفهارس `1` و `2`. بشكل عام، يمكننا إيجاد عقدة الأصل لأي عقدة عند أي فهرس بالمعادلة التالية: `Math.floor((i - 1) / 2)`. هذه الأنماط ستظل صحيحة مع نمو الشجرة الثنائية إلى أي حجم. أخيرًا، يمكننا إجراء تعديل بسيط لتسهيل هذه الحسابات أكثر بتخطي العنصر الأول في المصفوفة. يؤدي هذا إلى إنشاء العلاقة التالية لأي عنصر عند فهرس معين `i`:

تمثيل مصفوفي نموذجي:

```js
[ null, 6, 22, 30, 37, 63, 48, 42, 76 ]
```

العقدة الفرعية اليسرى للعنصر: `i * 2`

العقدة الفرعية اليمنى للعنصر: `i * 2 + 1`

عقدة الأصل للعنصر: `Math.floor(i / 2)`

بمجرد أن تفهم الحسابات، يصبح استخدام التمثيل المصفوفي مفيدًا جدًا لأن مواقع العقد يمكن تحديدها بسرعة بهذه الحسابات، ويقل استخدام الذاكرة لأنك لا تحتاج إلى الاحتفاظ بمراجع للعقد الفرعية.

# --instructions--

التعليمات: هنا سننشئ حزمة قصوى. ابدأ بإنشاء طريقة `insert` التي تضيف عناصر إلى الحزمة. أثناء الإدخال، من المهم دائمًا الحفاظ على خاصية الحزمة. بالنسبة للحزمة القصوى، هذا يعني أن العنصر الجذري يجب أن يكون دائمًا أكبر قيمة في الشجرة ويجب أن تكون جميع عقد الأصل أكبر من أبنائها. في تنفيذ الحزمة باستخدام مصفوفة، يتم تحقيق ذلك عادةً في ثلاث خطوات:

<ol>
<li>أضف العنصر الجديد إلى نهاية المصفوفة.</li>
  <li>إذا كان العنصر أكبر من عقدة الأصل، قم بتبديلهما.</li>
  <li>استمر في التبديل حتى يصبح العنصر الجديد أصغر من عقدة الأصل أو تصل إلى جذر الشجرة.</li>
</ol>

أخيرًا، أضف طريقة `print` التي تُرجع مصفوفة بكل العناصر التي أُضيفت إلى الحزمة.

# --hints--

يجب أن يوجد هيكل بيانات `MaxHeap`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    }
    return typeof test == 'object';
  })()
);
```

يجب أن يحتوي `MaxHeap` على طريقة تسمى `insert`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

يجب أن يحتوي `MaxHeap` على طريقة تسمى `print`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

يجب أن تضيف طريقة `insert` العناصر وفقًا لخاصية الحزمة القصوى.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    test.insert(50);
    test.insert(100);
    test.insert(700);
    test.insert(32);
    test.insert(51);
    test.insert(800);
    const result = test.print();
    const solution = JSON.stringify([null,800,51,700,32,50,100]);
    const solutionWithoutNull = JSON.stringify([800,51,700,32,50,100]);

    return (result.length == 6) ? (JSON.stringify(result) == solutionWithoutNull) : (JSON.stringify(result) == solution);
  })()
);
```

# --seed--

## --seed-contents--

```js
var MaxHeap = function() {
  // Only change code below this line
    
  // Only change code above this line
};
```

# --solutions--

```js
var MaxHeap = function() {
    // Only change code below this line
    this.heap = [];
    this.parent = index => {
      return Math.floor((index - 1) / 2);
    }
    this.insert = element => {
      this.heap.push(element);
      this.heapifyUp(this.heap.length - 1);
    }
    this.heapifyUp = index => {
      let currentIndex = index,
      parentIndex = this.parent(currentIndex);
      while (currentIndex > 0 && this.heap[currentIndex] > this.heap[parentIndex]) {
        this.swap(currentIndex, parentIndex);
        currentIndex = parentIndex;
        parentIndex = this.parent(parentIndex);
      }
    }
    this.swap = (index1, index2) => {
      [this.heap[index1], this.heap[index2]] = [this.heap[index2], this.heap[index1]];
    }
    this.print = () => {
      return this.heap;
    }
    // Only change code above this line
};
```
