---
id: 5cc0c1b32479e176caf3b422
title: التحقق مما إذا كانت الشجرة شجرة بحث ثنائي
challengeType: 1
forumTopicId: 301624
dashedName: check-if-tree-is-binary-search-tree
---

# --description--

بما أنك تعرف بالفعل ما هي شجرة البحث الثنائي، فإن هذا التحدٍّ سيحدد كيف يمكنك معرفة ما إذا كانت الشجرة شجرة بحث ثنائي أم لا.

الفرق الرئيسي في شجرة البحث الثنائي هو أن العقد مرتبة بطريقة منظمة. تحتوي العقد على عقدتين فرعيتين كحد أقصى (موضوعتين إلى اليمين و/أو اليسار) بناءً على ما إذا كانت قيمة العقدة الفرعية أكبر من أو تساوي (إلى اليمين) أو أقل من (إلى اليسار) عقدة الأصل.

# --instructions--

في هذا التحدٍّ، ستنشئ أداة لشجرتك. اكتب طريقة JavaScript باسم `isBinarySearchTree` تأخذ شجرة كمُدخَل وتُرجع قيمة منطقية تحدد ما إذا كانت الشجرة شجرة بحث ثنائي أم لا. استخدم العودية كلما أمكن.

# --hints--

يجب أن تُرجع شجرة البحث الثنائي true عند التحقق باستخدام `isBinarySearchTree()`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    test.push(1);
    test.push(5);
    test.push(3);
    test.push(2);
    test.push(4);
    return isBinarySearchTree(test) == true;
  })()
);
```

يجب أن تُرجع `isBinarySearchTree()` false عند التحقق من شجرة ليست شجرة بحث ثنائي.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    test.push(1);
    test.root.left = new Node(1);
    return isBinarySearchTree(test) == false;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype.push = function(val) {
  var root = this.root;

  if (!root) {
    this.root = new Node(val);
    return;
  }

  var currentNode = root;
  var newNode = new Node(val);

  while (currentNode) {
    if (val < currentNode.value) {
      if (!currentNode.left) {
        currentNode.left = newNode;
        break;
      } else {
        currentNode = currentNode.left;
      }
    } else {
      if (!currentNode.right) {
        currentNode.right = newNode;
        break;
      } else {
        currentNode = currentNode.right;
      }
    }
  }
};
```

# --seed--

## --seed-contents--

```js
var displayTree = (tree) => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
}
function isBinarySearchTree(tree) {
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = (tree) => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
}
function isBinarySearchTree(tree) {
  if (tree.root == null) {
    return null;
  } else {
    let isBST = true;
    function checkTree(node) {
      if (node.left != null) {
        const left = node.left;
        if (left.value >= node.value) {
          isBST = false;
        } else {
          checkTree(left);
        }
      }
      if (node.right != null) {
        const right = node.right;
        if (right.value < node.value) {
          isBST = false;
        } else {
          checkTree(right);
        }
      }
    }
    checkTree(tree.root);
    return isBST;
  }
};
```
