---
id: 6ac09f8420b4b67e075d2cd2
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このワークショップでは、最初のPythonプログラムを作成します。

ここがコードエディターで、Pythonのコードを書きます。1行目がハイライトされています。そこに移動して入力してください。

```py
'Hello'
```

行の先頭から始めて、引用符を忘れないようにしてください。

入力が終わったら、「check your code」ボタンをクリックして正しいか確認してください。

# --hints--

`'Hello'` と入力してください。行の先頭から始めて、引用符を忘れないようにしてください。

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""(['"])\\s*hello\\s*\\1"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
