---
id: 6ac0aaba914d81b32567c076
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

`print` は端末に出力を送るために使える組み込み関数です。ご覧の通り、`Hello` というテキストが端末に表示されています。

次に、ハイライトされた行に移動して、テキスト `freeCodeCamp` を端末に出力してください。

```py
print('freeCodeCamp')
```

行の先頭から始めるようにしてください。

# --hints--

コードには `print('freeCodeCamp')` があるはずです。行の先頭から始めるようにしてください。

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*freecodecamp\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
--fcc-editable-region--

--fcc-editable-region--
```
