---
id: 6ac0ae563f6de70930f47ef7
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

コードを書くとき、既存の値を再利用したり参照したりしたいことがよくあります。変数を使うとそれができます。

変数に値を代入するには、変数名の後に代入演算子（`=`）を書き、その変数に代入したい値を続けて書きます。

```py
sentence = 'I am learning Python'
```

ハイライトされている行に移動し、`sentence` という名前の変数を作成して、文字列 `I am learning Python` を代入してください。

# --hints--

コードには `sentence = 'I am learning Python'` があるはずです。行の先頭から書き始めていることを確認してください。

```js
({ test: () => runPython(`
import re
target_var = _Node(_code).find_variable("sentence")
pattern = r"""sentence\\s*=\\s*(['"])\\s*i\\s+am\\s+learning\\s+python\\s*\\1\\s*"""
assert re.fullmatch(pattern, str(target_var).strip(), flags=re.IGNORECASE)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')
--fcc-editable-region--

--fcc-editable-region--
```
