---
id: 6ac0b35f3b29a27f7e207295
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

前のステップで、変数 `sentence` を作成しました。次に、強調表示された行に `print(sentence)` と入力して、その値をターミナルに出力してください。

これで、最初のPythonプログラムが完成しました。

# --hints--

コードには `print(sentence)` が含まれている必要があります。行の先頭から始めていることを確認してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(sentence)")`)) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
print(sentence)
```
