---
id: 6ac0a4463662d9738c393e21
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

前のステップでは、**文字列**を作成しました。これは、シングルクォートまたは二重引用符で囲まれた文字のシーケンスです。プログラミングでは、文字列はテキストを表すために使われます。

**ターミナル**は、コードの出力を確認できる場所です。

出力を確認できるように、既存のコードの行を次のように変更してください。

```py
print('Hello')
```

行の先頭から始めるようにしてください。

# --hints--

コードには `print('Hello')` が含まれている必要があります。行の先頭から始めるようにしてください。

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*hello\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
'Hello'
--fcc-editable-region--
```
