---
id: 688140eac18479dee90553e5
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Kwa kuwa unatumia kipengele kingine cha `div`, utahitaji kufanya hili liwe na upatikanaji zaidi kwa wasomaji wa skrini.

Katika funzo la awali, ulijifunza kuhusu sifa ya kipengele ya `aria-live` ambayo hutumika kuunda eneo hai kwenye ukurasa. Hii itawawezesha watumiaji wa kisomaji cha skrini kupokea taarifa moja kwa moja wakati yaliyomo kwenye eneo hai yanapobadilika, bila ya kuhitaji kuweka makini au kuingiliana nayo kwa mikono.

Katika kipengele chako cha `#status`, ongeza sifa ya kipengele ya `aria-live` na weka thamani yake kuwa `"polite"`.

Thamani ya `polite` itawaambia wasomaji wa skrini kusubiri hadi tangazo lolote la sasa litakaposikika au hadi mtumizi aache kuandika kabla ya kutangaza sasisho hilo.

# --hints--

Kipengele chako cha `#status` kinapaswa kuwa na sifa ya kipengele ya `aria-live`.

```js
assert.isTrue(document.getElementById("status").hasAttribute("aria-live"));
```

Kipengele chako cha `#status` kinapaswa kuwa na sifa ya kipengele ya `aria-live` iliyowekwa kuwa `"polite"`.

```js
assert.equal(document.getElementById("status").getAttribute("aria-live"), "polite");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    <div id="note" class="note" contenteditable="true" aria-label="Note editor">
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English. 
        
      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.
        
      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>

    --fcc-editable-region--
    <div id="status"></div>
    --fcc-editable-region--

    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```
