---
id: 668f08ea07b99b1f4a91acab
title: Eine Rezeptseite erstellen
challengeType: 25
dashedName: build-a-recipe-page
demoType: onClick
---

# --description--

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine `!DOCTYPE html` Deklaration haben.
1. Sie sollten ein `html` Element mit `lang` gesetzt auf `en` haben.
1. Sie sollten ein `head` Element haben, das ein `title` Element mit dem Namen Ihres Rezepts enthält, und ein `meta` Element mit einem `charset` Attribut, das auf `UTF-8` gesetzt ist.
1. Sie sollten ein `body` Element haben.
1. Sie sollten ein `h1` Element mit dem Namen Ihres Rezepts haben.
1. Sie sollten ein `p` Element haben, das das Rezept unterhalb des `h1` einführt.
1. Sie sollten ein `h2` Element mit dem Text `Ingredients` für den Zutatenabschnitt haben.
1. Sie sollten eine ungeordnete Liste (`ul` Element) mit mindestens vier Listenelementen (`li` Elemente) haben, die Ihre Zutaten unterhalb des ersten `h2` Elements auflistet.
1. Sie sollten ein zweites `h2` Element mit dem Text `Instructions` für den Anweisungsabschnitt haben.
1. Sie sollten eine geordnete Liste (`ol` Element) mit mindestens vier Listenelementen haben, die die Rezeptschritte in der Reihenfolge unterhalb des zweiten `h2` auflistet.
1. Sie sollten ein `img` Element mit einem `src` Attribut haben, das auf ein gültiges Bild gesetzt ist. Sie können `https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg` verwenden, wenn Sie möchten, und ein `alt` Attribut, das das Bild beschreibt.

# --hints--

Ihre Rezeptseite sollte eine `!DOCTYPE html` Deklaration haben.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Sie sollten ein `html` Element mit `lang` gesetzt auf `en` haben.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/gi);
```

Sie sollten ein `head` Element innerhalb des `html` Elements haben.

```js
assert.match(code, /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i);
```

Sie sollten ein `title` Element innerhalb Ihres `head` Elements haben.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i);
```

Ihr `title` Element sollte Ihren Rezepttitel enthalten.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Sie sollten ein `meta` Element innerhalb Ihres `head` Elements haben.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*meta[\s\S]*>[\s\S]*<\/\s*head\s*>/i);
```

Ihr `meta` Element sollte sein `charset` Attribut auf `UTF-8` gesetzt haben.

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

Sie sollten ein `body` Element innerhalb Ihres `html` Elements haben.

```js
assert.match(code, /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i);
```

Sie sollten ein `h1` Element mit dem Namen Ihres Rezepts haben.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Sie sollten nur ein `h1` Element haben.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Sie sollten ein `p` Element unterhalb Ihres `h1` Elements haben.

```js
assert.strictEqual(document.querySelector('h1')?.nextElementSibling, document.querySelector('p'));
```

Ihr erstes `p` Element sollte Ihr Rezept beschreiben.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Ihr erstes `h2` Element sollte den Text `Ingredients` haben.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Ingredients');
```

Sie sollten ein ungeordnetes Listen-Element unterhalb Ihres ersten `h2` Elements haben.

```js
assert.strictEqual(document.querySelector('ul')?.previousElementSibling.tagName, 'H2');
```

Sie sollten mindestens vier Listenelemente in Ihrer ungeordneten Liste mit den Zutaten haben.

```js
const els = document.querySelectorAll('ul > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

Ihr zweites `h2` Element sollte den Text `Instructions` haben.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Instructions');
```

Sie sollten ein geordnetes Listen-Element unterhalb Ihres zweiten `h2` Elements haben.

```js
assert.strictEqual(document.querySelectorAll('h2')?.[1]?.nextElementSibling?.tagName, "OL");
```

Sie sollten mindestens vier Listenelemente in Ihrer geordneten Liste mit den Anweisungen haben.

```js
const els = document.querySelectorAll('ol > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

Sie sollten mindestens ein `img` Element haben.

```js
assert.exists(document.querySelector('img'));
```

Alle Ihre `img` Elemente sollten ein gültiges `src` Attribut und einen Wert haben.

```js
const img = document.querySelector('img');
const rawSrc = img?.getAttribute('src');
const resolvedSrc = img?.src;
const re = new RegExp(window.location.href, "ig");

assert.isAbove(rawSrc?.trim().length, 0, "The 'src' attribute must be explicitly set.");
assert.notMatch(resolvedSrc, re, "The 'src' should not start with the current page URL");

img.onload = () => {
  console.log('Image loaded successfully.');
};

img.onerror = (error) => {
  console.error('Image failed to load:', error);
  assert.fail("Your image's URL should be valid."); // Make the test instafail
};

if (img.complete) {
  img.onload && img.onload();
};
```

Alle Ihre `img` Elemente sollten ein `alt` Attribut haben, das das Bild beschreibt.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Chocolate chip cookies recipe</title>
</head>

<body>
  <h1>Chocolate Chip Cookies</h1>
  <p>Welcome to the ultimate guide for making mini chocolate chip cookies! These bite-sized treats are perfect for
    satisfying your sweet tooth without overindulging. Follow this simple recipe to create delicious,
    crispy-on-the-outside, chewy-on-the-inside mini chocolate chip cookies that everyone will love.</p>
  <img src="https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg" alt="Ingredients for baking: three eggs, a bowl of flour, a glass of milk, and a whisk arranged on a wooden table.">
  <h2>Ingredients</h2>
  <ul>
    <li>1 cup all-purpose flour</li>
    <li>1/2 teaspoon baking soda</li>
    <li>1/4 cup unsalted butter, softened</li>
    <li>1/4 cup granulated sugar</li>
    <li>1/2 teaspoon vanilla extract</li>
    <li>1/2 cup mini chocolate chips</li>
  </ul>
  <h2>Instructions</h2>
  <ol>
    <li>Preheat your oven to 350°F (175°C) and line a baking sheet with parchment paper.</li>
    <li>In a bowl, whisk together the flour and baking soda.</li>
    <li>In another bowl, beat the butter, sugar, and vanilla extract until creamy.</li>
    <li>Gradually add the dry ingredients to the wet mixture, then fold in the mini chocolate chips.</li>
    <li>Drop small spoonfuls of dough onto the baking sheet.</li>
    <li>Bake for 8-10 minutes, then let cool before enjoying!</li>
  </ol>
</body>

</html>
```
