---
id: a26cbbe9ad8655a977e1ceb5
title: بناء تطبيق للعثور على أطول كلمة
challengeType: 26
dashedName: build-a-longest-word-finder-app
---

# --description--

في هذا المختبر، ستبني دالة تُرجع طول أطول كلمة في الجملة المقدمة.

على سبيل المثال، في الجملة `"The quick brown fox jumped over the lazy dog"`، أطول كلمة هي `"jumped"`، وطولها 6.

**الهدف:** حقق قصص المستخدم أدناه واجتز جميع الاختبارات لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ دالة باسم `findLongestWordLength` تأخذ سلسلة نصية كمعلمة.
2. يجب أن تُرجع الدالة طول أطول كلمة في السلسلة النصية.

# --hints--

يجب أن تنشئ دالة باسم `findLongestWordLength`.

```js
assert.isFunction(findLongestWordLength);
```

يجب أن تحتوي `findLongestWordLength` على معلمة واحدة فقط.

```js
assert.lengthOf(findLongestWordLength, 1);
```

يجب أن تُرجع `findLongestWordLength("The quick brown fox jumped over the lazy dog")` رقمًا.

```js
assert.isNumber(
  findLongestWordLength('The quick brown fox jumped over the lazy dog')
);
```

يجب أن تُرجع `findLongestWordLength("The quick brown fox jumped over the lazy dog")` القيمة `6`.

```js
assert.strictEqual(
  findLongestWordLength('The quick brown fox jumped over the lazy dog'),
  6
);
```

يجب أن تُرجع `findLongestWordLength("May the force be with you")` القيمة `5`.

```js
assert.strictEqual(findLongestWordLength('May the force be with you'), 5);
```

يجب أن تُرجع `findLongestWordLength("Google do a barrel roll")` القيمة `6`.

```js
assert.strictEqual(findLongestWordLength('Google do a barrel roll'), 6);
```

يجب أن تُرجع `findLongestWordLength("Googling do a barrel roll")` القيمة `8`.

```js
assert.strictEqual(findLongestWordLength('Googling do a barrel roll'), 8);
```

يجب أن تُرجع `findLongestWordLength("What is the average airspeed velocity of an unladen swallow")` القيمة `8`.

```js
assert.strictEqual(
  findLongestWordLength(
    'What is the average airspeed velocity of an unladen swallow'
  ),
  8
);
```

يجب أن تُرجع `findLongestWordLength("What if we try a super-long word such as otorhinolaryngology")` القيمة `19`.

```js
assert.strictEqual(
  findLongestWordLength(
    'What if we try a super-long word such as otorhinolaryngology'
  ),
  19
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findLongestWordLength(str) {
  return str.split(' ').sort((a, b) => b.length - a.length)[0].length;
}

findLongestWordLength('The quick brown fox jumped over the lazy dog');
```
