---
id: 688140eac18479dee90553e5
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

Dato che stai usando un altro elemento `div`, dovrai renderlo più accessibile per i lettori di schermo.

In una lezione precedente, hai imparato l'attributo `aria-live` che si usa per creare una regione live su una pagina. Questo permette agli utenti di lettori di schermo di essere notificati automaticamente quando il contenuto della regione live cambia, senza dover mettere manualmente il focus o interagire con essa.

Nel tuo elemento `#status`, aggiungi un attributo `aria-live` e imposta il suo valore su `"polite"`.

Il valore `polite` dirà ai lettori di schermo di aspettare che qualsiasi annuncio in corso sia terminato o che l'utente smetta di digitare prima di annunciare l'aggiornamento.

# --hints--

Il tuo elemento `#status` dovrebbe avere l'attributo `aria-live`.

```js
assert.isTrue(document.getElementById("status").hasAttribute("aria-live"));
```

Il tuo elemento `#status` dovrebbe avere l'attributo `aria-live` impostato su `"polite"`.

```js
assert.equal(document.getElementById("status").getAttribute("aria-live"), "polite");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    <div id="note" class="note" contenteditable="true" aria-label="Note editor">
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English. 
        
      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.
        
      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>

    --fcc-editable-region--
    <div id="status"></div>
    --fcc-editable-region--

    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```
