---
id: 6626b4c58c027d86478ff5eb
title: 步驟 3
challengeType: 1
dashedName: step-3
---

# --description--

老師對你到目前為止創建的程式感到非常滿意。但現在他們想要一個簡單的方法來檢查學生是否及格。及格成績是任何不是 `"F"` 的成績。

完成函式（程式）`hasPassingGrade`，該函式（程式）以學生分數作為參數。若學生有及格成績，你的函式（程式）應傳回 `true`，否則傳回 `false`。

**提示**

- 使用 `getGrade` 函式（程式）來取得學生的成績。然後檢查該成績是否及格。

# --hints--

你的 `hasPassingGrade` 函式應該傳回一個布林值。

```js
assert.strictEqual(typeof hasPassingGrade(100), "boolean");
```

你的 `hasPassingGrade` 函式（程式）如果成績是 `"A"`，應該傳回 `true`。

```js
assert.isTrue(hasPassingGrade(100));
```

如果成績是 `"F"`，你的 `hasPassingGrade` 函式應該傳回 `false`。

```js
assert.isFalse(hasPassingGrade(53));
```

你的 `hasPassingGrade` 函式（程式）應該為所有及格成績傳回 `true`。

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

--fcc-editable-region--
function hasPassingGrade(score) {
  
}


console.log(hasPassingGrade(100));
console.log(hasPassingGrade(53));
console.log(hasPassingGrade(87));
--fcc-editable-region--
```
