---
id: 6626a060c4006f793e10cb33
title: 步驟 2
challengeType: 1
dashedName: step-2
---

# --description--

現在老師需要你幫忙將學生分數轉換成字母等級。

完成接受數字 `score` 作為參數的 `getGrade` 函式。你的函式應該傳回一個字串，表述根據分數的等級字母。

以下是分數及其對應的字母等級：

| Score Range | Grade   |
| ----------- | ------- |
| `100`       | `"A++"` |
| `90 - 99`   | `"A"`   |
| `80 - 89`   | `"B"`   |
| `70 - 79`   | `"C"`   |
| `60 - 69`   | `"D"`   |
| `0 - 59`    | `"F"`   |


**提示**

- 請記住你學過條件述語（`if`、`else if` 和 `else`）。
- 請記得你學過比較運算子（`>`、`<`、`>=`、`<=`、`===`）。
- 請記住你學過邏輯 AND 運算子（`&&`）。

# --hints--

你的 `getGrade` 函式（程式）應該在分數為 `100` 時傳回 `"A++"`。

```js
assert.strictEqual(getGrade(100), "A++");
```

如果分數是 `94`，你的 `getGrade` 函式（程式）應該傳回 `"A"`。

```js
assert.strictEqual(getGrade(94), "A");
```

如果分數介於 `80` 和 `89` 之間，你的 `getGrade` 函式（程式）應該傳回 `"B"`。

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

如果分數是 `78`，你的 `getGrade` 函式（程式）應該傳回 `"C"`。

```js
assert.strictEqual(getGrade(75), "C");
```

如果分數介於 `60` 和 `69` 之間，你的 `getGrade` 函式（程式）應該傳回 `"D"`。

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

如果分數是 `57`，你的 `getGrade` 函式（程式）應該傳回 `"F"`。

```js
assert.strictEqual(getGrade(57), "F");
```

如果分數介於 `0` 和 `59` 之間，你的 `getGrade` 函式（程式）應該傳回 `"F"`。

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}
--fcc-editable-region--
function getGrade(score) {

}

console.log(getGrade(96));
console.log(getGrade(82));
console.log(getGrade(56));
--fcc-editable-region--
```
