---
id: 6626b8dcf5057f896f948440
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

現在老師已經擁有所有需要的資訊，他們想要能夠將結果訊息傳送給學生。 

完成 `studentMsg` 函式（程式），並以 `totalScores` 和 `studentScore` 作為參數。該函式應傳回一個表述給學生的訊息字串。

如果學生通過課程，字串應該遵循此格式：

```md
Class average: average-goes-here. Your grade: grade-goes-here. You passed the course.
```

如果學生未通過課程，該字串應遵循此格式：

```md
Class average: average-goes-here. Your grade: grade-goes-here. You failed the course.
```

將 `average-goes-here` 取代為總分的平均值。  
將 `grade-goes-here` 取代為學生的成績。

**提示**

- 使用 `getAverage` 函式（程式）來取得類別平均值。
- 使用 `getGrade` 函式（程式）來取得學生的成績。
- 使用字串串接（`+`）來建置訊息。
- 請注意訊息中的標點符號和空格。

# --hints--

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` 應該傳回以下訊息：「類別平均：71.7。你的成績：F。你未通過課程。」。

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` 應該傳回以下訊息：「Class average: 50.8. Your grade: A++. You passed the course.」。

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A++. You passed the course.");
```

你的 `studentMsg` 函式應該根據學生的分數和類別平均值傳回正確的訊息。

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

--fcc-editable-region--
function studentMsg(totalScores, studentScore) {

}
console.log(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37));
--fcc-editable-region--

```


# --solutions--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

function studentMsg(totalScores, studentScore) {
  let average = getAverage(totalScores);
  let grade = getGrade(studentScore);

  return `Class average: ${average}. Your grade: ${grade}. You ${
    hasPassingGrade(studentScore) ? "passed" : "failed"
  } the course.`;
}

```
