---
id: 662693f82c91a66be46c881b
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

老師已經批改完學生的考試，並需要你的幫助來計算該類別的平均分數。

完成 `getAverage` 函式（程式），該函式（程式）接收一個測驗分數的陣列並傳回平均分數。

平均值是透過將所有分數加總後除以分數的總數來計算的。

```md
average = sum of all scores / total number of scores
```

已為你提供了幾個函式呼叫，以便你測試你的程式碼。

**提示**

- 你可以使用迴圈來迭代 `scores` 陣列並加總所有分數。
- 你可以使用 `length` 屬性來取得分數的總數。

# --hints--

你的 `getAverage` 函式應該傳回一個數字。

```js
assert.strictEqual(typeof getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 'number');
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` 應該傳回 `71.7`。

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` 應該傳回 `85.4`。

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` 應該傳回 `92.4`。

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

你的 `getAverage` 函式（程式）應該傳回測驗分數的平均分數。

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function getAverage(scores) {

}

console.log(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
console.log(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]));
--fcc-editable-region--
```
