---
id: 691f7bbf76229bb97c60827d
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop üben Sie die Verwendung von `for...of`-Schleifen, indem Sie eine Funktion erstellen, die zählt, wie oft ein String in einem Array von Strings vorkommt.

Bevor Sie diese Funktion schreiben, erstellen Sie eine einfachere, die jeden Buchstaben eines Strings in der Konsole ausgibt.

Beginnen Sie damit, eine leere Funktion namens `printCharacters` mit dem Parameter `str` zu definieren.

# --hints--

Sie sollten eine `printCharacters`-Funktion erstellen.

```js
assert.isFunction(printCharacters);
```

Ihre `printCharacters`-Funktion sollte einen Parameter namens `str` haben.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
