---
id: 6920c9cda1c7f056f75e4b1a
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Jetzt ist es Zeit, Ihre erste `for...of`-Schleife zu schreiben. In den vorherigen Lektionen haben Sie gelernt, dass Sie eine `for...of`-Schleife verwenden können, um Werte aus einem iterierbaren Objekt (z. B. Strings oder Arrays) zu durchlaufen.

Hier ist ein Beispiel für eine `for...of`-Schleife:

```js
for (const num of [1, 2, 3]) {
  // code block to be executed
}
```

Fügen Sie eine `for...of`-Schleife mit einem leeren Codeblock in Ihre Funktion ein. Sie soll über jedes Zeichen im `str`-Argument iterieren.

# --hints--

Ihre `printCharacters`-Funktion sollte eine `for...of`-Schleife in ihrem Codeblock enthalten.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functionprintCharacters\(str\)\{for\(\S+of\S+\)\{\S*?\}\}/);
```

Ihre `for...of`-Schleife sollte über das `str`-Argument Ihrer Funktion iterieren.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function printCharacters\(str\) \{\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+str/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
