---
id: 6920d189e36b6b60a6330fee
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Geben Sie nun innerhalb der Schleife die Variable `char` in der Konsole aus.

# --hints--

Sie sollten eine `console.log`-Anweisung mit `char` als Argument innerhalb der Schleife haben.

```js
const explorer = await __helpers.Explorer(code);
const normalizedFunctionCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(explorer.allFunctions.printCharacters?.toString()));
const loopCodeBlock = normalizedFunctionCode.match(/for\(constcharofstr\)\{(\S*?)}/)?.[1];
assert.match(loopCodeBlock, /console\.log\(char\)/);
```

Wenn die `printCharacters`-Funktion mit dem Argument `'test'` aufgerufen wird, sollte die `console.log()`-Anweisung innerhalb der Funktion viermal mit den folgenden Argumenten in dieser Reihenfolge aufgerufen werden: `'t'`, `'e'`, `'s'` und `'t'`. Stattdessen wurde sie mit den folgenden Argumenten aufgerufen: `--fcc-actual--`.

```js
const consoleLogSpy = __helpers.spyOn(console, 'log');
printCharacters('test');
assert.deepEqual(consoleLogSpy.calls, [['t'], ['e'], ['s'], ['t']]);
consoleLogSpy.restore();
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
