---
id: 6921d688b8d58581d34459b9
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

In den nächsten Schritten bauen Sie eine Funktion, die zählt, wie oft ein String in einem Array von Strings vorkommt.

Definieren Sie zunächst eine leere Funktion mit dem Namen `getMatchedWordCount` und den Parametern `sentence` und `match` in dieser Reihenfolge.

# --hints--

Sie sollten eine `getMatchedWordCount`-Funktion erstellen.

```js
assert.isFunction(getMatchedWordCount);
```

Ihre `getMatchedWordCount`-Funktion sollte zwei Parameter haben: `sentence` und `match`, in dieser Reihenfolge.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.getMatchedWordCount;
assert.lengthOf(parameters, 2);
assert.equal(parameters[0].toString(), 'sentence');
assert.equal(parameters[1].toString(), 'match');
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
