---
id: 6921de8d2a111d0341cff702
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Super, Ihre Funktion gibt ein `0` zurück, aber sie zählt noch nichts.

Um das zu beheben, erstellen Sie innerhalb von `getMatchedWordCount` eine `for...of`-Schleife mit einem leeren Codeblock, die jedes Wort in `sentence` durchläuft.

# --hints--

Ihre `getMatchedWordCount`-Funktion sollte eine `for...of`-Schleife innerhalb ihres Codeblocks enthalten.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functiongetMatchedWordCount\(sentence,match\)\{letcount=0\;for\(\S+of\S+\)\{\S*?\}returncount;\}/);
```

Ihre `for...of`-Schleife sollte das `sentence`-Argument Ihrer Funktion durchlaufen.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function getMatchedWordCount\(sentence, match\) \{\s*let count = 0\;\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*return count\;\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+sentence/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;

--fcc-editable-region--
  
--fcc-editable-region--

  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
