---
id: 6921d90b5ce50185505d2004
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Da die `getMatchedWordCount` Funktion eine numerische Zählung zurückgeben muss, beginnen Sie damit, eine Zählervariable einzurichten.

Erstellen Sie eine Variable namens `count` mit dem Schlüsselwort `let`, initialisieren Sie sie mit dem Wert `0` und fügen Sie eine `return`-Anweisung hinzu, die die `count` Variable zurückgibt.

# --hints--

Sie sollten die `count` Variable innerhalb der Funktion `getMatchedWordCount` deklarieren.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getMatchedWordCount?.variables.count);
```

Um die `count` Variable später inkrementieren zu können, deklarieren Sie sie mit dem Schlüsselwort `let`.

```js
const explorer = await __helpers.Explorer(code);
const countVar = explorer.allFunctions.getMatchedWordCount?.variables.count;
assert.isTrue(countVar?.toString().startsWith('let'));
```

Ihre Funktion sollte eine Zuweisung des Werts `0` an die `count` Variable enthalten.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getMatchedWordCount?.variables.count?.value.matches('0'));
```

Ihre Funktion sollte eine `return`-Anweisung enthalten, die die `count` Variable zurückgibt.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getMatchedWordCount?.hasReturn('count'));
```

Ihre `getMatchedWordCount` Funktion sollte den Wert `0` zurückgeben.

```js
assert.strictEqual(getMatchedWordCount(["foo", "bar"], "foo"), 0);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
