---
id: 6921dae10700e300f26994e3
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Um die Ausgabe Ihrer `getMatchedWordCount`-Funktion in ihrem aktuellen Zustand zu sehen, fügen Sie darunter eine `console.log`-Anweisung hinzu. Rufen Sie innerhalb der `console.log()` die Funktion mit dem Argument `["I", "really", "really", "really", "like", "to", "code"]` für den `sentence`-Parameter und `"really"` für den `match`-Parameter auf.

# --hints--

Sie sollten die Funktion `getMatchedWordCount` innerhalb einer `console.log`-Anweisung aufrufen.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
assert.match(normalizedCode, /console\.log\(getMatchedWordCount\(\S*?\)\)/);
```

Sie sollten die Funktion `getMatchedWordCount` mit dem Argument `["I", "really", "really", "really", "like", "to", "code"]` für den `sentence`-Parameter und `"really"` für den `match`-Parameter aufrufen.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/)[0];
assert.match(normalizedConsoleLogCode, /\[('|"|`)I\1,('|"|`)really\2,('|"|`)really\3,('|"|`)really\4,('|"|`)like\5,('|"|`)to\6,('|"|`)code\7\],('|"|`)really\8/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
