---
id: 6921e10ae2681205b516fbe6
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Jetzt ist es Zeit, die Logik in Ihre Schleife einzufügen, die `count` entsprechend erhöht.

Verwenden Sie innerhalb der Schleife eine bedingte Anweisung, um `count` um `1` zu erhöhen, wenn die Variable `word` gleich der Variable `match` ist. Andernfalls soll der Wert von `count` unverändert bleiben.

# --hints--

Das Aufrufen von `getMatchedWordCount([], "bar")`, `getMatchedWordCount(["foo"], "bar")`, `getMatchedWordCount(["foo", "bar", "foo"], "bar")` und `getMatchedWordCount(["foo", "bar", "foo"], "foo")` in dieser Reihenfolge sollte jeweils einen Wert vom Typ `"number"` zurückgeben; stattdessen wurden Werte vom Typ `--fcc-actual--` zurückgegeben.

```js
const obj = { method: (arg1, arg2) => typeof getMatchedWordCount(arg1, arg2) };
const spy = __helpers.spyOn(obj, 'method');
obj.method([], "bar");
obj.method(["foo"], "bar");
obj.method(["foo", "bar", "foo"], "bar");
obj.method(["foo", "bar", "foo"], "foo");
assert.deepEqual(spy.returns, ["number", "number", "number", "number"]);
spy.restore();
```

`getMatchedWordCount([], "bar")` sollte `0` zurückgeben, gibt aber `--fcc-actual--` zurück.

```js
const resultString = String(getMatchedWordCount([], "bar"));
assert.strictEqual(resultString, "0");
```

`getMatchedWordCount(["foo"], "bar")` sollte `0` zurückgeben, gibt aber `--fcc-actual--` zurück.

```js
const resultString = String(getMatchedWordCount(["foo"], "bar"));
assert.equal(resultString, "0");
```

`getMatchedWordCount(["foo", "bar", "foo"], "bar")` sollte `1` zurückgeben, gibt aber `--fcc-actual--` zurück.

```js
const resultString = String(getMatchedWordCount(["foo", "bar", "foo"], "bar"));
assert.equal(resultString, "1");
```

`getMatchedWordCount(["foo", "bar", "foo"], "foo")` sollte `2` zurückgeben, gibt aber `--fcc-actual--` zurück.

```js
const resultString = String(getMatchedWordCount(["foo", "bar", "foo"], "foo"));
assert.equal(resultString, "2");
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
--fcc-editable-region--
    
--fcc-editable-region--
    console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }
  
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
