---
id: 69220a03729c43082ad3e9f6
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

Sie haben Ihre `getMatchedWordCount`-Funktion fertiggestellt!

Jetzt werden Sie sie mit verschiedenen Daten testen, um zu sehen, wie sie sich verhﾃ､lt.

Fﾃｼgen Sie unter der vorhandenen eine neue `console.log`-Anweisung hinzu, die `getMatchedWordCount` mit diesen Parametern aufruft:

- sentence: `["Do", "not", "fear", "the", "dandy", "lion"]`
- match: `"dandy"`

Herzlichen Glﾃｼckwunsch! Sie haben diesen Workshop abgeschlossen.

# --hints--

Sie sollten eine zweite `console.log`-Anweisung haben, die die `getMatchedWordCount`-Funktion darin aufruft.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/g)[1];
assert.exists(normalizedConsoleLogCode);
```

Rufen Sie in Ihrer zweiten `console.log`-Anweisung die `getMatchedWordCount`-Funktion mit `["Do", "not", "fear", "the", "dandy", "lion"]` als `sentence`-Argument und `"dandy"` als `match`-Argument auf.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/g)[1];
assert.match(normalizedConsoleLogCode, /\[('|"|`)Do\1,('|"|`)not\2,('|"|`)fear\3,('|"|`)the\4,('|"|`)dandy\5,('|"|`)lion\6\],('|"|`)dandy\7/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
    if (word === match) {
      count++;
    }
    console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }
  
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
    if (word === match) {
      count++;
    }
    console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }
  
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);

console.log(
  getMatchedWordCount(
    ["Do", "not", "fear", "the", "dandy", "lion"],
    "dandy"
  )
);
```
