---
id: 6978bd465213832659f3338b
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

Geben Sie nun innerhalb der Schleife den folgenden Template-String in der Konsole aus: `Checking "${word}" against "${match}" | Running count: ${count}`

# --hints--

Sie sollten genau eine `console.log`-Anweisung innerhalb der Schleife in Ihrer `getMatchedWordCount`-Funktion haben.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const functionRegex = /(functiongetMatchedWordCount\(sentence,match\)\{\S*?returncount;\})/;
const normalizedFunctionCode = normalizedCode.match(functionRegex)[1];
console.log(normalizedFunctionCode);

assert.match(normalizedFunctionCode, /for\(constwordofsentence\)\{\S*?console\.log\(\S*?\)\S*?\}/);
assert.notMatch(normalizedFunctionCode, /for\(constwordofsentence\)\{\S*?(console\.log\(\S*?\))\S*?(console\.log\(\S*?\))\S*?\}/);
```

Ihre `console.log`-Anweisung innerhalb der Schleife in Ihrer `getMatchedWordCount`-Funktion sollte einen Template-String mit den folgenden Platzhaltern verwenden: `${word}`, `${match}` und `${count}`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const functionRegex = /(functiongetMatchedWordCount\(sentence,match\)\{\S*?returncount;\})/;
const normalizedFunctionCode = normalizedCode.match(functionRegex)[1];
const consoleLogArguments = normalizedFunctionCode.match(/for\(constwordofsentence\)\{\S*?console\.log\((\S*?)\)\S*?\}/)[1]
const templateLiteral = consoleLogArguments.match(/(^\`\S*\`$)/)[1];

assert.include(templateLiteral, "${word}");
assert.include(templateLiteral, "${match}");
assert.include(templateLiteral, "${count}");
```

Wenn die `getMatchedWordCount`-Funktion mit `(['foo'], 'bar')` aufgerufen wird, sollte die `console.log()`-Anweisung innerhalb der Funktion einmal mit dem String `Checking "foo" against "bar" | Running count: 0` aufgerufen werden. Stattdessen wurde sie mit jedem String in der folgenden durch Kommas getrennten Liste aufgerufen: `--fcc-actual--`.

```js
const consoleLogSpy = __helpers.spyOn(console, 'log');
getMatchedWordCount(['foo'], 'bar');
assert.deepEqual(consoleLogSpy.calls, [['Checking "foo" against "bar" | Running count: 0']]);
consoleLogSpy.restore();
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
--fcc-editable-region--
    
--fcc-editable-region--
  }

  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
