---
id: 699d20623e57e8ffee88545a
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

카탈로그를 파싱했으니, 필드별로 검색할 수 있습니다. `findByAuthor`는 작성자에 검색어가 포함된 항목을 필터링할 것입니다.

`findByAuthor`와 `catalog`를 매개변수로 받는 `author`이라는 함수를 만드세요. 그 안에서 `searchTerm`의 소문자 버전으로 `author`라는 변수를 선언해 나중에 대소문자 구분 없이 일치시킬 수 있게 하세요.

또한 일치하는 항목을 담을 빈 배열을 `results`라는 변수에 할당해 선언하세요.

# --hints--

`findByAuthor`와 `catalog`를 매개변수로 받는 `author`라는 함수가 있어야 합니다.

```js
const regex = __helpers.functionRegex('findByAuthor', ['catalog', 'author']);
assert.match(__helpers.removeJSComments(code), regex);
```

`searchTerm`로 `author.toLowerCase()`라는 변수를 선언해야 합니다.

```js
assert.match(findByAuthor.toString(), /searchTerm\s*=\s*author\.toLowerCase\s*\(\s*\)/);
```

빈 배열로 `results`라는 변수를 선언해야 합니다.

```js
assert.match(findByAuthor.toString(), /(const|let|var)\s+results\s*=\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

function parseCatalog(rawCards) {
  const catalog = [];
  for (let i = 0; i < rawCards.length; i++) {
    catalog.push(parseCard(rawCards[i]));
  }
  return catalog;
}

const catalog = parseCatalog(rawCatalogCards);
--fcc-editable-region--

--fcc-editable-region--
```
