---
id: 699d20623e57e8ffee88545e
title: 20단계
challengeType: 1
dashedName: step-20
---

# --description--

반복문이 완료되었습니다. 함수에서 `grouped`을 반환하세요.

그다음, 함수 아래에서 `groupByDecade(catalog)`을 호출하고 결과를 `byDecade`라는 변수에 할당하세요. `byDecade`을 기록하여 그룹화된 결과를 확인하면 `"1950s"`, `"1960s"`, `"1970s"` 같은 키가 있는 객체가 보이고, 각 키는 해당 10년대의 책 배열을 가리키고 있어야 합니다.

# --hints--

`groupByDecade`은 객체를 반환해야 합니다.

```js
assert.isObject(groupByDecade(catalog));
```

`byDecade["1980s"].length`은 `5`이어야 합니다.

```js
assert.strictEqual(groupByDecade(catalog)["1980s"].length, 5);
```

`byDecade["Unknown"].length`은 `1`이어야 합니다.

```js
assert.strictEqual(groupByDecade(catalog)["Unknown"].length, 1);
```

`groupByDecade(catalog)`의 결과를 `byDecade`이라는 변수에 할당해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+byDecade\s*=\s*groupByDecade\s*\(\s*catalog\s*\)/);
```

`byDecade`을 콘솔에 기록해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*byDecade\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

function parseCatalog(rawCards) {
  const catalog = [];
  for (let i = 0; i < rawCards.length; i++) {
    catalog.push(parseCard(rawCards[i]));
  }
  return catalog;
}

const catalog = parseCatalog(rawCatalogCards);

function findByAuthor(catalog, author) {
  const searchTerm = author.toLowerCase();
  const results = [];
  for (let i = 0; i < catalog.length; i++) {
    if (catalog[i].author.toLowerCase().includes(searchTerm)) {
      results.push(catalog[i]);
    }
  }
  return results;
}

function groupByDecade(catalog) {
  const grouped = {};
  for (let i = 0; i < catalog.length; i++) {
    const book = catalog[i];
    if (book.year === "Unknown") {
      if (!grouped["Unknown"]) {
        grouped["Unknown"] = [];
      }
      grouped["Unknown"].push(book);
      continue;
    }
    const decade = Math.floor(book.year / 10) * 10;
    const decadeKey = `${decade}s`;
    if (!grouped[decadeKey]) {
      grouped[decadeKey] = [];
    }
    grouped[decadeKey].push(book);
  }
--fcc-editable-region--

}
--fcc-editable-region--
```
