---
id: 699d20623e57e8ffee885455
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

콘솔 출력에서 볼 수 있듯이, 분할한 후에도 각 부분에는 여전히 앞뒤에 불필요한 공백이 있습니다. 각 요소를 개별적으로 처리해서 모두 정리해야 합니다.

이럴 때 `for` 반복문이 완벽한 사용 사례입니다. 반복문은 배열 같은 반복 가능한 객체의 각 요소마다 코드 블록을 한 번씩 실행합니다.

```js
for (let i = 0; i < array.length; i++) {
  // code to run on each iteration
}
```

`i`은 `0` (첫 번째 인덱스)에서 시작하고, `i`이 배열 길이보다 작을 동안 반복하며, `i++`는 각 반복 후 `i`를 1씩 증가시킵니다.

`parseCard` 안에서 `trimmedParts`라는 빈 배열을 선언하세요. 그리고 `for` 반복문을 추가해 `parts`를 순회하세요. 반복문 본문은 아직 비워두고 다음 단계에서 채울 예정입니다.

# --hints--

빈 배열로 설정된 `trimmedParts`라는 변수를 선언해야 합니다.

```js
assert.match(parseCard.toString(), /(const|let|var)\s+trimmedParts\s*=\s*\[\s*\]/);
```

`for`를 순회하기 위해 `parts` 반복문을 사용해야 합니다.

```js
assert.match(parseCard.toString(), /for\s*\(\s*(var|let)\s+(\w+)\s*=\s*0\s*;\s*\2\s*<\s*parts\s*\.\s*length\s*;\s*\2\+\+\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
--fcc-editable-region--

--fcc-editable-region--
  return parts;
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
