---
id: 699d20623e57e8ffee885456
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

이제 부품들이 정리되었으니, 이를 구조화된 객체로 변환할 수 있습니다. JavaScript에서 객체는 중괄호와 "키: 값" 쌍을 사용합니다. 예를 들어, `{ title: "Dune", year: 1965 }`.

네 개의 변수를 선언하세요: `title`, `author`, `year`, 그리고 `location`. 각각에 `trimmedParts`의 해당 요소를 할당하세요. 그런 다음 `return` 문을 업데이트하여 이 네 개의 속성을 가진 객체를 반환하게 하세요.

# --hints--

`parseCard`은 `title`, `author`, `year`, 그리고 `location` 속성을 가진 객체를 반환해야 합니다.

```js
const result = parseCard("The Stand | King, Stephen | 1978 | Shelf K2");
assert.isObject(result);
assert.property(result, "title");
assert.property(result, "author");
assert.property(result, "year");
assert.property(result, "location");
```

`title`이 반환하는 객체의 `parseCard("The Stand | King, Stephen | 1978 | Shelf K2")` 속성은 `"The Stand"` 값을 가져야 합니다.

```js
assert.strictEqual(parseCard("The Stand | King, Stephen | 1978 | Shelf K2").title, "The Stand");
```

`author`이 반환하는 객체의 `parseCard("The Stand | King, Stephen | 1978 | Shelf K2")` 속성은 `"King, Stephen"` 값을 가져야 합니다.

```js
assert.strictEqual(parseCard("The Stand | King, Stephen | 1978 | Shelf K2").author, "King, Stephen");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
--fcc-editable-region--

  return trimmedParts;
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
