---
id: 699d20623e57e8ffee885450
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

나머지 카탈로그 카드가 추가되었습니다. 끝부분 근처에 필드가 누락된 세 개의 항목이 있으니, 나중 단계에서 처리할 것입니다.

`parseCard`는 원시 카탈로그 문자열을 받아 네 개의 프로퍼티 `title`, `author`, `year`, `location`를 가진 구조화된 객체를 반환합니다. 지금은 `parseCard`라는 매개변수를 가진 `rawString` 함수를 만드세요.

# --hints--

`parseCard`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(parseCard);
```

`parseCard`는 `rawString`라는 매개변수를 가진 함수여야 합니다.

```js
const regex = __helpers.functionRegex('parseCard', ['rawString']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

--fcc-editable-region--

--fcc-editable-region--
```
