---
id: 699d20623e57e8ffee885457
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

일부 카탈로그 카드에는 누락된 필드가 있습니다. 예를 들어, 16번 항목은 작성자가 비어 있어서 `parseCard`가 현재 해당 필드에 빈 문자열을 반환합니다.

`||` 연산자는 대체 수단을 제공합니다: `a || b`는 `a`가 참일 때 `a`로 평가되고, 그렇지 않으면 `b`로 평가됩니다. 빈 문자열은 거짓이므로 `"" || "Unknown"`는 `"Unknown"`로 평가됩니다.

return 문에서 `title`, `author`, `location`을 업데이트하여 `|| "Unknown"`을 대체 수단으로 사용하세요. `year`는 다음 단계에 남겨두세요.

# --hints--

작성자 필드가 비어 있는 문자열로 호출할 때, 반환된 객체의 `author` 속성은 `"Unknown"`여야 합니다.

```js
assert.strictEqual(parseCard("X |  | 2000 | Y").author, "Unknown");
```

제목 필드가 비어 있는 문자열로 호출할 때, 반환된 객체의 `title` 속성은 `"Unknown"`여야 합니다.

```js
assert.strictEqual(parseCard(" | A | 2000 | Y").title, "Unknown");
```

위치 필드가 비어 있는 문자열로 호출할 때, 반환된 객체의 `location` 속성은 `"Unknown"`여야 합니다.

```js
assert.strictEqual(parseCard("X | A | 2000 | ").location, "Unknown");
```

아직 `year`에 대한 대체 수단을 추가하지 마세요.

```js
assert.strictEqual(parseCard("X | A |  | Y").year, "");
```

모든 필드가 있을 때, `parseCard`은 실제 값을 반환해야 하며 `"Unknown"`이 아니어야 합니다.

```js
assert.strictEqual(parseCard("The Stand | King, Stephen | 1978 | Shelf K2").title, "The Stand");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
--fcc-editable-region--
  return {
    title: title,
    author: author,
    year: year,
    location: location
  };
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
