---
id: 699d20623e57e8ffee885459
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

`parseCard`는 테스트를 거쳐 제대로 작동합니다. `cardResult` 변수는 더 이상 필요하지 않습니다.

먼저 `const cardResult` 줄과 `console.log(cardResult)` 줄을 삭제하세요. 그런 다음 `parseCatalog`를 매개변수로 받는 `rawCards`라는 함수를 만드세요. 지금은 함수 내부에 빈 배열 `catalog`를 선언하고 반환하세요.

함수 아래에서 `parseCatalog(rawCatalogCards)`를 호출하고, 결과를 `catalog`라는 변수에 할당한 뒤 `catalog.length`를 콘솔에 출력하세요.

# --hints--

`parseCatalog`는 배열을 반환하는 함수여야 합니다.

```js
assert.isArray(parseCatalog(rawCatalogCards));
```

`catalog`는 선언되어 `parseCatalog(rawCatalogCards)`로 설정되어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+catalog\s*=\s*parseCatalog\s*\(/);
```

`cardResult` 변수와 관련된 `console.log()` 호출을 삭제해야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /\bcardResult\b/);
```

`catalog.length`를 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*catalog\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

--fcc-editable-region--
const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
--fcc-editable-region--
```
