---
id: 699d20623e57e8ffee885451
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

`parseCard` 함수는 원시 카탈로그 문자열을 받아 구조화된 객체를 반환합니다. 함수를 작성하기 전에 작업 중에 출력 결과를 확인할 방법을 설정하세요.

`parseCard`을 `rawCatalogCards[2]` 인수로 호출하고, 결과를 `cardResult`라는 변수에 할당한 뒤 `cardResult`를 콘솔에 출력하세요.

현재 `parseCard`는 본문이 없어서 `undefined`를 반환합니다. 다음 단계에서 함수에 코드를 추가하면, 출력된 결과를 통해 각 단계에서 함수가 반환하는 값을 정확히 확인할 수 있습니다.

# --hints--

`parseCard`을 `rawCatalogCards[2]` 인수로 호출해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /parseCard\s*\(\s*rawCatalogCards\s*\[\s*2\s*\]\s*\)/);
```

`parseCard`을 `rawCatalogCards[2]` 인수로 호출한 결과를 `cardResult`라는 변수에 할당해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+cardResult\s*=\s*parseCard\s*\(\s*rawCatalogCards\s*\[\s*2\s*\]\s*\)\s*;?/);
```

`cardResult`를 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*cardResult\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {}

--fcc-editable-region--

--fcc-editable-region--
```
