---
id: 699d20623e57e8ffee885454
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

카탈로그 카드를 구문 분석하는 첫 번째 단계는 네 개의 필드를 분리하는 것입니다. `.split()` 메서드는 제공한 구분자가 나타날 때마다 문자열을 하위 문자열 배열로 나눕니다.

예를 들어:

```js
"The Shining | King, Stephen | 1977 | Shelf K1".split("|")
// ["The Shining ", " King, Stephen ", " 1977 ", " Shelf K1"]
```

`parseCard` 안에서 `.split("|")`에 `rawString`를 호출하고 결과를 반환하세요. 이 단계를 마치면 콘솔에 네 개의 문자열로 이루어진 배열이 표시되어야 합니다.

**참고:** 값 주변에 여분의 공백이 있을 수 있습니다. 다음 단계에서 이를 정리할 것입니다.

# --hints--

`.split("|")` 함수 내에서 `rawString`에 `parseCard`를 호출해야 합니다.

```js
assert.match(parseCard.toString(), /rawString\s*\.\s*split\s*\(\s*('|"|`)\|\1\s*\)\s*;?/);
```

`parseCard`는 배열을 반환해야 합니다.

```js
assert.isArray(parseCard("A | B | C | D"));
```

`"A | B | C | D"`를 `parseCard`와 함께 호출하면 네 개 요소가 있는 배열을 반환해야 합니다.

```js
assert.strictEqual(parseCard("A | B | C | D").length, 4);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
--fcc-editable-region--
  
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
