---
id: 699d1f8e51fcdc8230f779fb
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 유산 도서관 목록 시스템을 구축합니다. 원시 도서 데이터를 파싱하고, 항목을 검색 및 그룹화하며, 데이터 품질을 검증하고, 결과를 다양한 형식으로 내보냅니다.

각 목록 카드는 `"Title | Author | Year | Location"` 형식의 단일 문자열로 저장되며, 필드는 파이프 문자(`|`)로 구분됩니다.

워크숍 전반에 걸쳐 이러한 문자열 배열을 사용하여 목록 함수를 만들고 테스트합니다. `rawCatalogCards`라는 배열을 만들고 다음 두 문자열을 추가하세요:

```js
"From a Buick 8 | King, Stephen | 2002 | Shelf K7",
"The Shining | King, Stephen | 1977 | Shelf K1"
```

# --hints--

`rawCatalogCards`라는 변수를 선언하세요.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+rawCatalogCards\s*=/);
```

`rawCatalogCards`는 배열이어야 합니다.

```js
assert.isArray(rawCatalogCards);
```

`rawCatalogCards`의 첫 번째 요소는 `"From a Buick 8 | King, Stephen | 2002 | Shelf K7"`이어야 합니다.

```js
assert.strictEqual(rawCatalogCards[0], "From a Buick 8 | King, Stephen | 2002 | Shelf K7");
```

`rawCatalogCards`의 두 번째 요소는 `"The Shining | King, Stephen | 1977 | Shelf K1"`이어야 합니다.

```js
assert.strictEqual(rawCatalogCards[1], "The Shining | King, Stephen | 1977 | Shelf K1");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
