---
id: 699d20623e57e8ffee88545d
title: 19단계
challengeType: 1
dashedName: step-19
---

# --description--

`"Unknown"` 확인 후, 연도가 숫자인 책들을 10년 단위로 그룹화하는 로직을 추가하세요. `const decade = Math.floor(book.year / 10) * 10`을 선언하세요. 예를 들어, `Math.floor(1977 / 10)`는 `197`를 주고, `10`를 곱하면 `1970`가 됩니다.

그다음 `const decadeKey =` `` `${decade}s` `` to produce a string like `"1970s"`를 선언하세요.

책을 그룹화할 때는 대괄호 표기법을 사용하세요: `grouped[decadeKey]`이 아직 존재하지 않는지 확인하고, 비어 있는 배열로 초기화한 다음 `book`을 그 안에 추가하세요. 키가 변수에 저장되어 있으므로 여기서는 대괄호 표기법이 필요합니다.

# --hints--

`decade`라는 변수를 선언하고 `Math.floor(book.year / 10) * 10`으로 설정하세요.

```js
assert.match(groupByDecade.toString(), /(const|let|var)\s+decade\s*=\s*Math\.floor\s*\(\s*book\.year\s*\/\s*10\s*\)\s*\*\s*10/);
```

`decadeKey`라는 변수를 선언하고 `decade`와 문자 `"s"`를 포함하세요.

```js
assert.match(groupByDecade.toString(), /(const|let|var)\s+decadeKey\s*=.*decade/s);
assert.match(groupByDecade.toString(), /decadeKey.*s/s);
```

`grouped[decadeKey]`이 아직 존재하지 않는지 확인하고 비어 있는 배열로 초기화하세요.

```js
assert.match(groupByDecade.toString(), /!\s*grouped\s*\[\s*decadeKey\s*\]/);
assert.match(groupByDecade.toString(), /grouped\s*\[\s*decadeKey\s*\]\s*=\s*\[\s*\]/);
```

`book`을 `grouped[decadeKey]`에 추가하세요.

```js
assert.match(groupByDecade.toString(), /grouped\s*\[\s*decadeKey\s*\]\s*\.push\s*\(\s*book\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

function parseCatalog(rawCards) {
  const catalog = [];
  for (let i = 0; i < rawCards.length; i++) {
    catalog.push(parseCard(rawCards[i]));
  }
  return catalog;
}

const catalog = parseCatalog(rawCatalogCards);

function findByAuthor(catalog, author) {
  const searchTerm = author.toLowerCase();
  const results = [];
  for (let i = 0; i < catalog.length; i++) {
    if (catalog[i].author.toLowerCase().includes(searchTerm)) {
      results.push(catalog[i]);
    }
  }
  return results;
}

function groupByDecade(catalog) {
  const grouped = {};
  for (let i = 0; i < catalog.length; i++) {
    const book = catalog[i];
    if (book.year === "Unknown") {
      if (!grouped["Unknown"]) {
        grouped["Unknown"] = [];
      }
      grouped["Unknown"].push(book);
      continue;
    }
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
