---
id: 699d20623e57e8ffee88545b
title: 15단계
challengeType: 1
dashedName: step-15
---

# --description--

검색이 제대로 작동하는지 확인해 보겠습니다. Stephen King의 책이 카탈로그에 여섯 권 있습니다.

`findByAuthor(catalog, "king")`를 호출하고 결과를 `kingBooks`라는 변수에 할당하세요. `kingBooks.length`를 기록하여 개수를 확인하세요.

그다음 `for` 반복문을 사용해 각 결과를 기록하세요. 템플릿 리터럴을 사용해 제목과 연도를 포함하세요. 예: `` `${kingBooks[i].title} (${kingBooks[i].year})` ``.

# --hints--

`kingBooks`의 길이는 `6`이어야 합니다.

```js
assert.strictEqual(kingBooks.length, 6);
```

`kingBooks.length`를 콘솔에 기록해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*kingBooks\.length\s*\)/);
```

각 책의 제목과 연도를 기록하려면 `for` 반복문을 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*(var|let)\s+(\w+)\s*=\s*0\s*;\s*\2\s*<\s*kingBooks\s*\.\s*length\s*;\s*\2\+\+\s*\)\s*\{/);
assert.match(__helpers.removeJSComments(code), /kingBooks\[.+\]\.title\b/);
assert.match(__helpers.removeJSComments(code), /kingBooks\[.+\]\.year\b/);
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

function parseCatalog(rawCards) {
  const catalog = [];
  for (let i = 0; i < rawCards.length; i++) {
    catalog.push(parseCard(rawCards[i]));
  }
  return catalog;
}

const catalog = parseCatalog(rawCatalogCards);

function findByAuthor(catalog, author) {
  const searchTerm = author.toLowerCase();
  const results = [];
  for (let i = 0; i < catalog.length; i++) {
    if (catalog[i].author.toLowerCase().includes(searchTerm)) {
      results.push(catalog[i]);
    }
  }
  return results;
}

--fcc-editable-region--

--fcc-editable-region--
```
