---
id: 69a65722a69bc6f8d87a9952
title: 26단계
challengeType: 1
dashedName: step-26
---

# --description--

JSON(JavaScript Object Notation)은 시스템 간 데이터를 공유하는 표준 형식입니다. `JSON.stringify`는 JavaScript 값을 JSON 형식 문자열로 변환합니다.

테스트용이었던 두 개의 `validateEntry` 콘솔.log 호출을 제거하는 것부터 시작하세요. 그런 다음 `exportToJSON`를 매개변수로 받는 `catalog`라는 함수를 만드세요. `JSON.stringify(catalog, null, 2)`를 반환하세요. 첫 번째 인수는 변환할 데이터입니다. 두 번째 인수(`null`)는 사용자 지정 필터링이 없음을 의미합니다. 세 번째 인수(`2`)는 두 칸 들여쓰기를 추가해 출력 결과를 읽기 쉽게 만듭니다.

함수 아래에 `exportToJSON(catalog.slice(0, 2))`를 로그로 출력해 처음 두 항목의 JSON 출력을 미리 확인하세요.

# --hints--

`validateEntry` 콘솔.log 호출을 제거해야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\s*\(\s*validateEntry/);
```

`exportToJSON`는 문자열을 반환해야 합니다.

```js
assert.isString(exportToJSON(catalog));
```

`JSON.parse(exportToJSON(catalog))`는 유효한 배열을 생성해야 합니다.

```js
assert.isArray(JSON.parse(exportToJSON(catalog)));
```

두 번째와 세 번째 인수로 `JSON.stringify`와 `null`를 사용해 `2`를 호출해야 합니다.

```js
assert.match(exportToJSON.toString(), /JSON\.stringify\s*\(\s*catalog\s*,\s*null\s*,\s*2\s*\)/);
```

`exportToJSON(catalog.slice(0, 2))`를 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*exportToJSON\s*\(/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

function parseCatalog(rawCards) {
  const catalog = [];
  for (let i = 0; i < rawCards.length; i++) {
    catalog.push(parseCard(rawCards[i]));
  }
  return catalog;
}

const catalog = parseCatalog(rawCatalogCards);

function findByAuthor(catalog, author) {
  const searchTerm = author.toLowerCase();
  const results = [];
  for (let i = 0; i < catalog.length; i++) {
    if (catalog[i].author.toLowerCase().includes(searchTerm)) {
      results.push(catalog[i]);
    }
  }
  return results;
}

function groupByDecade(catalog) {
  const grouped = {};
  for (let i = 0; i < catalog.length; i++) {
    const book = catalog[i];
    if (book.year === "Unknown") {
      if (!grouped["Unknown"]) {
        grouped["Unknown"] = [];
      }
      grouped["Unknown"].push(book);
      continue;
    }
    const decade = Math.floor(book.year / 10) * 10;
    const decadeKey = `${decade}s`;
    if (!grouped[decadeKey]) {
      grouped[decadeKey] = [];
    }
    grouped[decadeKey].push(book);
  }
  return grouped;
}

const byDecade = groupByDecade(catalog);

function renderEntry(entry) {
  const title = entry.title || "Unknown";
  const author = entry.author || "Unknown";
  const year = entry.year || "Unknown";
  const location = entry.location || "Unknown";
  return `${"-".repeat(25)}
Title: ${title}
Author: ${author}
Year: ${year}
Location: ${location}
${"-".repeat(25)}`;
}

console.log(renderEntry(catalog[0]));

function validateEntry(entry) {
  let isValid = true;
  if (!("title" in entry) || !entry.title || entry.title === "Unknown") {
    isValid = false;
  }
  if (!("author" in entry) || !entry.author || entry.author === "Unknown") {
    isValid = false;
  }
  if (!("year" in entry) || !entry.year || entry.year === "Unknown") {
    isValid = false;
  }
  if (!("location" in entry) || !entry.location || entry.location === "Unknown") {
    isValid = false;
  }
  return isValid;
}

--fcc-editable-region--
console.log(validateEntry(catalog[0]));
console.log(validateEntry(catalog[16]));

--fcc-editable-region--
```
