---
id: 69a9de7ceb9b18f4fc988306
title: 14단계
challengeType: 1
dashedName: step-14
---

# --description--

루프 안에서 현재 항목의 작성자(소문자로 변환한)가 `searchTerm`를 포함하는지 확인하세요. `.includes()` 메서드는 문자열이 다른 문자열을 부분 문자열로 포함하면 `true`를 반환합니다. 조건이 참이면 항목을 `results`에 추가하세요. 루프가 끝난 후 `results`를 반환하세요.

# --hints--

`findByAuthor(catalog, "king").length`는 `6`여야 합니다.

```js
assert.strictEqual(findByAuthor(catalog, "king").length, 6);
```

`findByAuthor(catalog, "asimov").length`는 `2`여야 합니다.

```js
assert.strictEqual(findByAuthor(catalog, "asimov").length, 2);
```

`findByAuthor(catalog, "xyz").length`는 `0`여야 합니다.

```js
assert.strictEqual(findByAuthor(catalog, "xyz").length, 0);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

function parseCatalog(rawCards) {
  const catalog = [];
  for (let i = 0; i < rawCards.length; i++) {
    catalog.push(parseCard(rawCards[i]));
  }
  return catalog;
}

const catalog = parseCatalog(rawCatalogCards);

function findByAuthor(catalog, author) {
  const searchTerm = author.toLowerCase();
  const results = [];
  for (let i = 0; i < catalog.length; i++) {
    --fcc-editable-region--

  }
--fcc-editable-region--
}
```
