---
id: 69aa11c2584523e6e991330c
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

`.trim()` 메서드는 문자열 양쪽 끝의 공백을 제거합니다. 예를 들어, `" hello ".trim()`은 `"hello"`를 반환합니다.

루프 본문 안에서 `parts[i].trim()`을 `trimmedParts`에 추가하세요. 그런 다음 기존 `return` 문장을 `trimmedParts` 대신 `parts`를 반환하도록 업데이트하세요.

이 단계를 완료하면, 기록된 출력에 불필요한 공백 없이 깔끔한 값이 표시되어야 합니다.

# --hints--

루프 안에서 `parts[i].trim()`을 `trimmedParts`에 추가해야 합니다.

```js
assert.match(parseCard.toString(), /trimmedParts\s*\.\s*push\s*\(\s*parts\s*\[\s*\w+\s*\]\s*\.\s*trim\s*\(\s*\)\s*\)/);
```

`parseCard`은 `trimmedParts`를 반환해야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { parseCard } = explorer.allFunctions;
assert.isTrue(parseCard.hasReturn("trimmedParts"));
```

`parseCard`은 각 요소가 앞뒤 공백이 없는 배열을 반환해야 합니다.

```js
assert.isArray(parseCard("  A  |  B  |  C  |  D  "));
```

`"  A  |  B  |  C  |  D  "`을 인수로 호출하면 반환된 배열의 첫 번째 요소는 `"A"`여야 합니다.

```js
assert.strictEqual(parseCard("  A  |  B  |  C  |  D  ")[0], "A");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
--fcc-editable-region--

  }
  return parts;
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
