---
id: 69d5e6f7a8b9c0d1e2f3a4b5
title: 22단계
challengeType: 1
dashedName: step-22
---

# --description--

이제 `renderEntry`를 채우세요. 각 필드에 대해 대체 수단 상수를 선언하고 `const title = entry.title || "Unknown"`에 적용하세요. `author`, `year`, `location`에도 같은 패턴을 적용하세요.

도서관 카드 형식으로 항목을 포맷하는 여러 줄 템플릿 리터럴을 반환하세요. 위아래 테두리에는 `"-".repeat(25)`를 사용하고, `Title`, `Author`, `Year`, `Location`에 대해 레이블이 붙은 줄을 포함하세요. 예를 들어, 제목에 대한 레이블 줄은 `Title: ${title}`처럼 보입니다.

# --hints--

`renderEntry`는 문자열을 반환해야 합니다.

```js
assert.isString(renderEntry(catalog[0]));
```

`catalog[0]`로 호출할 때 반환된 문자열에 `"From a Buick 8"`이 포함되어야 합니다.

```js
assert.include(renderEntry(catalog[0]), "From a Buick 8");
```

`catalog[0]`로 호출할 때 반환된 문자열에 `"King, Stephen"`가 포함되어야 합니다.

```js
assert.include(renderEntry(catalog[0]), "King, Stephen");
```

`catalog[0]`로 호출할 때 반환된 문자열에 연도 `2002`가 포함되어야 합니다.

```js
assert.include(renderEntry(catalog[0]), "2002");
```

`catalog[0]`로 호출할 때 반환된 문자열에 `"Shelf K7"`가 포함되어야 합니다.

```js
assert.include(renderEntry(catalog[0]), "Shelf K7");
```

반환된 문자열에는 `"-".repeat(25)`를 사용한 위아래 테두리가 포함되어야 합니다.

```js
assert.isAbove(renderEntry(catalog[0]).split("-".repeat(25)).length, 2);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

function parseCatalog(rawCards) {
  const catalog = [];
  for (let i = 0; i < rawCards.length; i++) {
    catalog.push(parseCard(rawCards[i]));
  }
  return catalog;
}

const catalog = parseCatalog(rawCatalogCards);

function findByAuthor(catalog, author) {
  const searchTerm = author.toLowerCase();
  const results = [];
  for (let i = 0; i < catalog.length; i++) {
    if (catalog[i].author.toLowerCase().includes(searchTerm)) {
      results.push(catalog[i]);
    }
  }
  return results;
}

function groupByDecade(catalog) {
  const grouped = {};
  for (let i = 0; i < catalog.length; i++) {
    const book = catalog[i];
    if (book.year === "Unknown") {
      if (!grouped["Unknown"]) {
        grouped["Unknown"] = [];
      }
      grouped["Unknown"].push(book);
      continue;
    }
    const decade = Math.floor(book.year / 10) * 10;
    const decadeKey = `${decade}s`;
    if (!grouped[decadeKey]) {
      grouped[decadeKey] = [];
    }
    grouped[decadeKey].push(book);
  }
  return grouped;
}

const byDecade = groupByDecade(catalog);

function renderEntry(entry) {
--fcc-editable-region--

--fcc-editable-region--
}

console.log(renderEntry(catalog[0]));
```
