---
id: 587d8259367417b2b2512c86
title: "implementare l'ordinamento per inserimento"
challengeType: 1
forumTopicId: 301613
dashedName: implement-insertion-sort
---

# --description--

Il prossimo metodo di ordinamento che vedremo è l'ordinamento per inserimento. Questo metodo funziona costruendo un array ordinato all'inizio della lista. Inizia l'array ordinato con il primo elemento. Poi esamina il prossimo elemento e lo scambia all'indietro nell'array ordinato finché non è nella posizione corretta. Continua a iterare attraverso la lista e a scambiare i nuovi elementi all'indietro nella parte ordinata finché non raggiunge la fine. Questo algoritmo ha una complessità temporale quadratica nei casi medi e peggiori.

**Istruzioni:** Scrivi una funzione `insertionSort` che prende un array di numeri interi come input e restituisce un array di questi numeri ordinati dal più piccolo al più grande.

# --hints--

`insertionSort` dovrebbe essere una funzione.

```js
assert.isFunction(insertionSort);
```

`insertionSort` dovrebbe restituire un array ordinato (dal più piccolo al più grande).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    insertionSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`insertionSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` dovrebbe restituire un array che non cambia tranne che per l'ordine.

```js
assert.sameMembers(
  insertionSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`insertionSort([5, 4, 33, 2, 8])` dovrebbe restituire `[2, 4, 5, 8, 33]`.

```js
assert.deepEqual(insertionSort([5, 4, 33, 2, 8]), [2, 4, 5, 8, 33])
```

`insertionSort` non dovrebbe usare il metodo integrato `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    insertionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function insertionSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function insertionSort (array) {
  for (let currentIndex = 0; currentIndex < array.length; currentIndex++) {
    let current = array[currentIndex];
    let j = currentIndex - 1;
    while (j > -1 && array[j] > current) {
      array[j + 1] = array[j];
      j--;
    }
    array[j + 1] = current;
  }
  return array;
}
```
