---
id: 587d825c367417b2b2512c8f
title: implementare merge sort
challengeType: 1
forumTopicId: 301614
dashedName: implement-merge-sort
---

# --description--

Un altro algoritmo di ordinamento intermedio comune è il merge sort. Come il quick sort, anche il merge sort usa una metodologia ricorsiva di divide-et-impera per ordinare un array. Sfrutta il fatto che è relativamente facile ordinare due array a patto che ciascuno sia già ordinato. Ma iniziamo con un solo array come input, quindi come facciamo a ottenere due array ordinati da quello? Possiamo dividere ricorsivamente l’input originale in due fino a raggiungere il caso base di un array con un solo elemento. Un array con un solo elemento è naturalmente ordinato, quindi possiamo iniziare a unire. Questa combinazione farà tornare indietro le chiamate ricorsive che hanno diviso l’array originale, producendo infine un array finale ordinato con tutti gli elementi. I passaggi del merge sort, quindi, sono:

**1)** Dividere ricorsivamente a metà l’array di input fino a produrre un sotto-array con un solo elemento.

**2)** Unire ciascun sotto-array ordinato per produrre l’array finale ordinato.

Il merge sort è un metodo di ordinamento efficiente, con complessità temporale di *O(nlog(n))*. Questo algoritmo è popolare perché è performante e relativamente semplice da implementare.

A margine, questo sarà l’ultimo algoritmo di ordinamento che tratteremo qui. Tuttavia, più avanti nella sezione sulle strutture dati ad albero descriveremo l’heap sort, un altro metodo di ordinamento efficiente che richiede un heap binario nella sua implementazione.

**Istruzioni:** Scrivi una funzione `mergeSort` che prende un array di numeri interi come input e restituisce un array di questi numeri ordinati dal più piccolo al più grande. Un buon modo per implementarla è scrivere una funzione, per esempio `merge`, che si occupa di unire due array ordinati, e un’altra funzione, per esempio `mergeSort`, che si occupa della ricorsione che produce array con un solo elemento da passare all’unione. Buona fortuna!

# --hints--

`mergeSort` dovrebbe essere una funzione.

```js
assert.isFunction(mergeSort);
```

`mergeSort` dovrebbe restituire un array ordinato (dal più piccolo al più grande).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    mergeSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`mergeSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` dovrebbe restituire un array che non cambia se non per l’ordine.

```js
assert.sameMembers(
  mergeSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`mergeSort` non dovrebbe usare il metodo integrato `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    mergeSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function mergeSort(array) {
  if (array.length === 1) {
    return array;
  } else {
    const splitIndex = Math.floor(array.length / 2);
    return merge(
      mergeSort(array.slice(0, splitIndex)),
      mergeSort(array.slice(splitIndex))
    );
  }

  // Merge two sorted arrays
  function merge(array1, array2) {
    let merged = [];
    while (array1.length && array2.length) {
      if (array1[0] < array2[0]) {
        merged.push(array1.shift());
      } else if (array1[0] > array2[0]) {
        merged.push(array2.shift());
      } else {
        merged.push(array1.shift(), array2.shift());
      }
    }

    // After looping ends, one array is empty, and other array contains only
    // values greater than all values in `merged`
    return [...merged, ...array1, ...array2];
  }
}

mergeSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]);
```
