---
id: 587d825a367417b2b2512c89
title: implementare quick sort
challengeType: 1
forumTopicId: 301615
dashedName: implement-quick-sort
---

# --description--

Qui passeremo a un algoritmo di ordinamento intermedio: quick sort. Quick sort è un metodo efficiente, ricorsivo e basato sulla strategia divide-et-impera per ordinare un array. In questo metodo, si sceglie un valore pivot nell'array originale. L'array viene quindi diviso in due sotto-array di valori minori e maggiori del valore pivot. Combiniamo poi il risultato della chiamata ricorsiva dell'algoritmo quick sort su entrambi i sotto-array. Questo processo continua fino a raggiungere il caso base di un array vuoto o con un solo elemento, che viene restituito. Lo srotolamento delle chiamate ricorsive ci restituisce l'array ordinato.

Quick sort è un metodo di ordinamento molto efficiente, che offre una prestazione media di *O(nlog(n))*. È anche relativamente semplice da implementare. Queste caratteristiche lo rendono un metodo di ordinamento popolare e utile.

**Istruzioni:** Scrivi una funzione `quickSort` che prende un array di numeri interi come input e restituisce un array di questi numeri ordinati dal più piccolo al più grande. Anche se la scelta del valore pivot è importante, per i nostri scopi qui qualsiasi pivot va bene. Per semplicità, si può usare il primo o l'ultimo elemento.

# --hints--

`quickSort` dovrebbe essere una funzione.

```js
assert(typeof quickSort == 'function');
```

`quickSort` dovrebbe restituire un array ordinato (dal più piccolo al più grande).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    quickSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`quickSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` dovrebbe restituire un array che non cambia se non nell'ordine.

```js
assert.sameMembers(
  quickSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quickSort` non dovrebbe usare il metodo integrato `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quickSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function quickSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function quickSort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quickSort(lesser), ...equal, ...quickSort(greater)];
  }
}
```
