---
id: 61abc7ebf3029b56226de5b6
title: implementare la ricerca binaria
challengeType: 1
forumTopicId: 487618
dashedName: implement-binary-search
---

# --description--

La ricerca binaria è un algoritmo con efficienza **O(log(n))** per cercare un elemento in un array ordinato. Funziona seguendo questi passaggi:

1. Trova l'elemento centrale `value` di un array ordinato. Se `value == target` restituisci `true` (Il valore è stato trovato e la ricerca è completa).
1. Se il valore centrale `value < target`, cerca nella metà destra dell'array nel confronto successivo.
1. Se il valore centrale `value > target`, cerca nella metà sinistra dell'array nel confronto successivo.
1. Se dopo aver cercato in tutto l'array il valore non è presente, restituisci `false` (L'array è stato cercato e il valore non è nell'array).

Come vedi, dimezzi successivamente un array, il che ti dà l'efficienza log(n). Per questa sfida, vogliamo che mostri il tuo lavoro — come sei arrivato al valore obiettivo... il percorso che hai seguito!

# --instructions--

Scrivi una funzione `binarySearch` che implementa l'algoritmo di ricerca binaria su un array, restituendo il percorso seguito (ogni valore centrale confrontato) per trovare l'obiettivo in un array.

La funzione prende come input un array ordinato di numeri interi e un valore obiettivo. Restituisce un array contenente (in ordine) il valore centrale trovato a ogni dimezzamento dell'array originale fino a trovare il valore obiettivo. Il valore obiettivo dovrebbe essere l'ultimo elemento dell'array restituito. Se il valore non viene trovato, restituisci la stringa `Value Not Found`.

Per esempio, `binarySearch([1,2,3,4,5,6,7], 5)` restituirebbe `[4,6,5]`.

Per questa sfida, quando dimezzi, DEVI usare `Math.floor()` durante la divisione: `Math.floor(x/2)`. Questo garantirà un percorso coerente e testabile.

**Nota:** L'array seguente sarà usato nei test:

```js
const testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
```

# --hints--

`binarySearch` dovrebbe essere una funzione.

```js
assert.isFunction(binarySearch);
```

`binarySearch(testArray, 0)` dovrebbe restituire `[13, 5, 2, 0]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 0), [13, 5, 2, 0]);
```

`binarySearch(testArray, 1)` dovrebbe restituire `[13, 5, 2, 0, 1]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 1), [13, 5, 2, 0, 1]);
```


`binarySearch(testArray, 2)` dovrebbe restituire `[13, 5, 2]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 2), [13, 5, 2]);
```

`binarySearch(testArray, 6)` dovrebbe restituire la stringa `Value Not Found`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.strictEqual(binarySearch(_testArray, 6), 'Value Not Found');
```

`binarySearch(testArray, 11)` dovrebbe restituire `[13, 5, 10, 11]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 11), [13, 5, 10, 11])
```

`binarySearch(testArray, 13)` dovrebbe restituire `[13]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 13), [13]);
```

`binarySearch(testArray, 70)` dovrebbe restituire `[13, 19, 22, 49, 70]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 70), [13, 19, 22, 49, 70]);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let arrayPath = [];
  return arrayPath;
}
```



# --solutions--

```js
let binarySearch = (searchList, value) => {
  let arrayPath = [];

  // set initial L - M - R
  let left = 0;
  let right = searchList.length - 1;
  let middle = Math.floor(right / 2);

  // if first comparison finds value
  if (searchList[middle] == value) {
    arrayPath.push(searchList[middle]);
    return arrayPath;
  }

  while (searchList[middle] !== value) {
    // add to output array
    arrayPath.push(searchList[middle]);

    // not found
    if (right < left) {
      return 'Value Not Found';
    }
    // value is in left or right portion of array
    // update L - M - R
    if (searchList[middle] > value) {
      right = middle - 1;
      middle = left + Math.floor((right - left) / 2);
    } else {
      left = middle + 1;
      middle = left + Math.floor((right - left) / 2);
    }

    // if found update output array and exit
    if (searchList[middle] == value) {
      arrayPath.push(searchList[middle]);

      break;
    }
  }
  return arrayPath;
};
```
