---
id: 6565a536ba1f9f25bd30e88b
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在此项目中，你将实现 Luhn 算法。 该算法是各种识别验证号码的公式。

首先声明一个名为 `main` 的函数，它将作为程序的入口点。 使用 `pass` 关键字以避免错误。

# --hints--

你的代码中应该有 `def main():` 。

```js
({
    test: () => {
        assert(runPython(`
        import inspect
        inspect.isfunction(main)
        `))
    }
})
```

你的代码中应该有 `pass` 。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const pass = __helpers.python.getDef("\n"+transformedCode, "main");
        assert.include(pass?.function_body, "pass");
    }
})
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
