---
id: 587d824a367417b2b2512c44
title: Verificador de preço de ações
challengeType: 4
forumTopicId: 301572
dashedName: stock-price-checker
---

# --description--

Construa um aplicativo JavaScript full stack que seja funcionalmente semelhante a este: <a href="https://stock-price-checker.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://stock-price-checker.freecodecamp.rocks/</a>.

Como todas as APIs de preço de ações confiáveis requerem uma chave de API, criamos uma solução alternativa.Use <a href="https://stock-price-checker-proxy.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://stock-price-checker-proxy.freecodecamp.rocks/</a> para obter informações atualizadas sobre preços de ações sem precisar se cadastrar para obter sua própria chave. 

Trabalhar nesse projeto vai fazer com que você escreva o código usando um dos seguintes métodos:

-   Clone <a href="https://github.com/freeCodeCamp/boilerplate-project-stockchecker/" target="_blank" rel="noopener noreferrer nofollow">este repositório do GitHub</a> e complete seu projeto localmente.
-   Use um construtor de sites de sua escolha para completar o projeto. Certifique-se de incorporar todos os arquivos do nosso repositório no GitHub.

# --instructions--

1.  Defina a variável de ambiente `NODE_ENV` como `test`, sem aspas
2.  Complete o projeto em `routes/api.js` ou criando um manipulador/controlador
3.  Você adicionará todas as funcionalidades de segurança ao `server.js`
4.  Você criará todos os testes funcionais em `tests/2_functional-tests.js`

**Nota** Considerações sobre Privacidade: Devido à exigência de que apenas 1 curtida por IP seja aceita, você terá que salvar os endereços IP.É importante manter a conformidade com as leis relativas à privacidade de dados, como o Regulamento Geral para a Proteção de Dados.Uma opção é obter permissão para salvar os dados do usuário, mas é muito mais simples deixar os dados anônimos.Para este desafio, lembre-se de anonimizar endereços IP antes de salvá-los no banco de dados.Se você precisa de ideias sobre como fazer isso, você pode optar por criptografar os dados (algoritmo hash), omitir parte deles, ou definir parte do endereço IP como 0.

Escreva os testes a seguir em `tests/2_functional-tests.js`:

-   Visualizar uma ação: faça a solicitação de GET para `/api/stock-prices/`
-   Visualizar uma ação e dizer que gostou dela: faça a solicitação de GET para `/api/stock-prices/`
-   Visualizar a mesma ação e dizer que gostou dela novamente: faça a solicitação de GET para `/api/stock-prices/`
-   Visualizar duas ações: faça a solicitação de GET para `/api/stock-prices/`
-   Visualizar duas ações e dizer que gostou delas: faça a solicitação de GET para `/api/stock-prices/`

# --hints--

Você pode fornecer seu próprio projeto, não o exemplo de URL.

```js
  assert(
    !/.*\/stock-price-checker\.freecodecamp\.rocks/.test(code)
  );
```

Você deve definir as políticas de segurança do conteúdo para permitir apenas o carregamento de scripts e CSS do seu servidor.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(
    parsed.headers['content-security-policy'].includes("script-src 'self'")
  );
  assert.isTrue(
    parsed.headers['content-security-policy'].includes("style-src 'self'")
  );
```

Você pode enviar uma solicitação de `GET` para `/api/stock-prices`, passando um símbolo da ação na NASDAQ para um parâmetro de consulta de `stock`.O objeto retornado conterá uma propriedade chamada `stockData`.

```js
  const data = await fetch(
    code + '/api/stock-prices?stock=GOOG'
  );
  const parsed = await data.json();
  assert.property(parsed, 'stockData');
```

A propriedade `stockData` inclui o símbolo de `stock` como uma string, o `price` como um número e `likes` como um número.

```js
  const data = await fetch(
    code + '/api/stock-prices?stock=GOOG'
  );
  const parsed = await data.json();
  const ticker = parsed.stockData;
  assert.typeOf(ticker.price, 'number');
  assert.typeOf(ticker.likes, 'number');
  assert.typeOf(ticker.stock, 'string');
```

Você também pode passar adiante um campo `like` como `true` (booleano) para ter sua curtida adicionada ao(s) estoque(s).Apenas 1 curtida por IP deve ser aceita.

```js

```

Se você passar 2 ações, o valor retornado será um array com informações sobre as duas.Em vez de `likes`, será exibido `rel_likes` (a diferença entre curtidas entre ambas as ações) para os dois objetos de `stockData`.

```js
  const data = await fetch(
    code + '/api/stock-prices?stock=GOOG&stock=MSFT'
  );
  const parsed = await data.json();
  const ticker = parsed.stockData;
  assert.typeOf(ticker, 'array');
  assert.property(ticker[0], 'rel_likes');
  assert.property(ticker[1], 'rel_likes');
```

Todos os 5 testes funcionais estão completos e passando.

```js
  const tests = await fetch(code + '/_api/get-tests');
  const parsed = await tests.json();
  assert.isTrue(parsed.length >= 5);
  parsed.forEach((test) => {
    assert.equal(test.state, 'passed');
  });
```

