---
id: 673b567e3ba535dda140d278
title: Jenga mfumo wa kupiga kura
challengeType: 26
dashedName: build-a-voting-system
saveSubmissionToDB: true
---

# --description--

Katika maabara hii, u-tajenga mfumo wa kupiga kura unaotumia `Map` kuunda kura na `Set` kuzuia upigaji kura mara mbili.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuanzisha kigezo cha `poll` kuwa kitu kipya cha `Map`.

2. Unapaswa kuwa na kitendakazi `addOption` kinachokubali kigezo `option`.

3. Katika kitendakazi cha `addOption`:

   - Ikiwa `option` bado haipo katika kura, inapaswa kuongezwa katika kura na `Set` tupu kama thamani yake kufuatilia wapigaji kura. Pia unapaswa kurudisha ujumbe `Option "<option>" added to the poll.`

   - Ikiwa `option` tayari ipo, inapaswa kurudisha ujumbe `Option "<option>" already exists.`

   - Ikiwa unajaribu kuongeza chaguo tupu, kitendakazi kinapaswa kurudisha ujumbe `Option cannot be empty.`

4. Unapaswa kuwa na kitendakazi `vote` kinachokubali vigezo viwili, `option` (chaguo la kupigia kura) na `voterId` (kitambulisho cha kipekee cha mpigaji kura).

5. Katika kitendakazi cha `vote`:

   - Ikiwa `option` haipo katika kura, kitendakazi kinapaswa kurudisha ujumbe `Option "<option>" does not exist.`

   - Ikiwa `option` ipo, kitendakazi kinapaswa kuangalia kama `voterId` tayari amepiga kura kwa `option` hii.

   - Ikiwa mpigaji kura tayari amepiga kura, kitendakazi kinapaswa kurudisha ujumbe `Voter <voterId> has already voted for "<option>".`

   - Ikiwa mpigaji kura hajapiga kura, `voterId` inapaswa kuongezwa kwenye `Set` ya `voters` kwa chaguo hili. Kitendakazi kinapaswa kurudisha ujumbe `Voter <voterId> voted for "<option>".`

6. Unapaswa kuwa na angalau chaguzi tatu katika `poll` yako.

7. `poll` yako inapaswa kuwa na angalau kura tatu.

8. Unapaswa kuwa na kitendakazi `displayResults` kinachorudisha matokeo ya kura kwa muundo ufuatao:

```js
Poll Results:
OptionA: N votes
OptionB: N votes
.
.

/*
sample output

Poll Results:
Turkey: 2 votes
Morocco: 1 votes
*/
```

# --hints--

Unapaswa kuwa na kigezo cha `poll` kilichoanzishwa kuwa kitu kipya cha `Map`.

```js
assert.exists(poll);
assert.instanceOf(poll, Map);
```

Unapaswa kufafanua kitendakazi `addOption` kinachokubali kigezo `option`.

```js
assert.isFunction(addOption);
assert.match(
  addOption.toString(),
  /^function\s+addOption\s*\(\s*option\s*\)\s*\{|\(?\s*option\s*\)?\s*=>/
);
```

Unapaswa kufafanua kitendakazi `vote` kinachokubali vigezo viwili, `option` na `voterId`.

```js
assert.isFunction(vote);
assert.match(vote.toString(), /\(\s*option\s*,\s*voterId\s*\)/);
```

Unapaswa kufafanua kitendakazi `displayResults` kuonyesha matokeo ya kura.

```js
assert.isFunction(displayResults);
```

Unapaswa kuwa na angalau chaguzi tatu katika `poll` yako.

```js
assert.isAtLeast(poll.size, 3);
```

`poll` yako inapaswa kuwa na angalau kura tatu.

```js
let totalVotes = 0;
for (let voters of poll.values()) {
  totalVotes += voters.size;
}
assert.isAtLeast(totalVotes, 3);
```

Unapaswa kuhakikisha kila chaguo la kupigia kura linaambatana na kitu cha `Set`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();

  addOption('Turkey');
  assert.instanceOf(poll.get('Turkey'), Set);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`addOption("Egypt")` inapaswa kurudisha `Option "Egypt" added to the poll.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(addOption('Egypt'), 'Option "Egypt" added to the poll.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Kuongeza chaguo tupu kunapaswa kurudisha `"Option cannot be empty."`

```js
assert.equal(addOption(''), 'Option cannot be empty.');
```

Wakati `Turkey` tayari imeongezwa, `addOption("Turkey")` inapaswa kurudisha `Option "Turkey" already exists.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Turkey');
  assert.equal(addOption('Turkey'), 'Option "Turkey" already exists.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Wakati `Malaysia` ipo katika chaguzi za kupigia kura, `vote("Malaysia", "traveler1")` inapaswa kurudisha `Voter traveler1 voted for "Malaysia".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Malaysia');
  assert.equal(
    vote('Malaysia', 'traveler1'),
    'Voter traveler1 voted for "Malaysia".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`vote` inapaswa kusasisha `Set` ya wapigaji kura kwa chaguo.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Alaska');
  vote('Alaska', 'seal');
  assert.isTrue(poll.get('Alaska').has('seal'));
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Wakati `traveler1` anajaribu kupiga kura kwa `Algeria` tena, `vote("Algeria", "traveler1")` inapaswa kurudisha `Voter traveler1 has already voted for "Algeria".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Algeria');
  vote('Algeria', 'traveler1');
  assert.equal(
    vote('Algeria', 'traveler1'),
    'Voter traveler1 has already voted for "Algeria".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Kura za mara mbili hazipaswi kuongeza ukubwa wa `Set`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  vote('Turkey', 'traveler1');
  const voterSet = poll.get('Turkey');
  const initialSize = voterSet.size;

  vote('Turkey', 'traveler1');
  assert.equal(voterSet.size, initialSize);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Wakati `Nigeria` haipo katika chaguzi za kupigia kura, `vote("Nigeria", "traveler2")` inapaswa kurudisha `Option "Nigeria" does not exist.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(
    vote('Nigeria', 'traveler2'),
    'Option "Nigeria" does not exist.'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Chaguo la kipekee linaweza kupokea kura nyingi.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Bali');
  assert.equal(vote('Bali', 'traveler1'), 'Voter traveler1 voted for "Bali".');
  assert.equal(vote('Bali', 'traveler2'), 'Voter traveler2 voted for "Bali".');
  assert.equal(poll.get('Bali').size, 2);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Wakati chaguo `Turkey` kinapata kura mbili, chaguo `Morocco` kinapata kura moja, na chaguo `Spain` hakipati kura yoyote, `displayResults()` inapaswa kurudisha `"Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes"`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');
  addOption('Spain');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`displayResults()` inapaswa kurudisha matokeo kwa muundo sahihi.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const poll = new Map();

function addOption(option) {
  if (!option || option.trim() === '') {
    return 'Option cannot be empty.';
  }
  if (!poll.has(option)) {
    poll.set(option, new Set());
    return `Option "${option}" added to the poll.`;
  } else {
    return `Option "${option}" already exists.`;
  }
}

function vote(option, voterId) {
  if (!poll.has(option)) {
    return `Option "${option}" does not exist.`;
  }
  const voters = poll.get(option);
  if (voters.has(voterId)) {
    return `Voter ${voterId} has already voted for "${option}".`;
  } else {
    voters.add(voterId);
    return `Voter ${voterId} voted for "${option}".`;
  }
}

function displayResults() {
  let results = 'Poll Results:\n';
  for (let [option, voters] of poll.entries()) {
    results += `${option}: ${voters.size} votes\n`;
  }
  return results.trim();
}

addOption('Turkey');
addOption('Morocco');
addOption('Spain');

vote('Turkey', 'traveler1');
vote('Turkey', 'traveler2');
vote('Morocco', 'traveler3');
```
