---
id: 69de4308e652aaeb5ccdb03e
title: 第 1 步
challengeType: 20
dashedName: step-1
---

# --description--

在本次工作坊中，你将练习使用 Python 的字串方法。

首先创建一个名为 `email` 的变量，并将字串 `alice.johnson@company.com` 赋值给它。

# --hints--

你应该创建一个名为 `email` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email')
`) })
```

你应该将字串 `alice.johnson@company.com` 赋值给 `email`。记得用单引号或双引号将文本括起来。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('email').is_equivalent("email = 'alice.johnson@company.com'")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
