---
id: 69e20876baed33de636b989a
title: 第 3 步
challengeType: 20
dashedName: step-3
---

# --description--

有时，你可能想查找字串中特定字符的位置。Python 提供了内置的 `.find()` 方法来实现这个目的。它返回一个整数，表示给定子字串首次出现的索引（位置）。

这是一个示例：

```py
text = 'hello'
print(text.find('e'))  # 1
```

现在，在 `email` 上调用 `.find()` 方法，并传入 `'@'`，以查找 `email` 字串中 @ 符号的位置。将结果赋值给名为 `at_position` 的变量。

然后打印 `at_position`，使输出显示为：`Position of @: 13`。

# --hints--

你应该创建一个名为 `at_position` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_position')
`) })
```

你应该将 `at_position` 赋值为 `email.find('@')` 的值。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_position').is_equivalent("at_position = email.find('@')")
`) })
```

你应该打印 `at_position`，使输出显示为 `Position of @: 13`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Position of @: {at_position}')",
    "print('Position of @: ' + str(at_position))",
    "print('Position of @:', at_position)",
    "print('Position of @:', str(at_position))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')
--fcc-editable-region--

--fcc-editable-region--
```
