---
id: 69e20a0451c7a3ef189514d3
title: 第 4 步
challengeType: 20
dashedName: step-4
---

# --description--

既然你知道了 `@` 在电子邮件中的位置，就可以提取放在 `@` 符号前面的用户名。

在之前的工作坊中，你使用切片提取了字串的一部分。作为提醒，语法是：

```python
text = 'hello world'
print(text[:5])  # hello
```

对 `email` 使用切片，提取 `at_position` 之前的所有内容，并将结果保存到名为 `username` 的变量中。

然后打印 `username`，使输出显示为 `Username: alice.johnson`。

# --hints--

你应该创建一个名为 `username` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('username')
`) })
```

你应该对 `email` 使用带有 `at_position` 的切片来赋值给 `username`。

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('username')
assert node.is_equivalent("username = email[:at_position]") or node.is_equivalent("username = email[0:at_position]")
`) })
```

你应该打印 `username`，使输出显示为 `Username: alice.johnson`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {username}')",
    "print('Username: ' + username)",
    "print('Username:', username)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

--fcc-editable-region--

--fcc-editable-region--
```
