---
id: 69e6299710319e7cb8ebdc37
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

真实世界的数据通常很混乱，带有多余的空格和不一致的大小写。

创建一个变量 `messy_name` 并赋值为带有两个前导空格和两个行尾空白的字串 `sARaH dAVis`。然后打印它，使输出显示为：

```md
Original name: '  sARaH dAVis  '
```

你将在接下来的步骤中清理它。

# --hints--

你应该创建一个名为 `messy_name` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_name')
`) })
```

你应该将 <code>'  sARaH dAVis  '</code> 赋值给 `messy_name`。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_name').is_equivalent("messy_name = '  sARaH dAVis  '")
`) })
```

你应该打印 `messy_name`，使输出显示为 `Original name: '  sARaH dAVis  '`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original name: '{messy_name}'\\")",
    "print(\\"Original name: '\\" + messy_name + \\"'\\")",
    "print(\\"Original name: '\\", messy_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

--fcc-editable-region--

--fcc-editable-region--
```
