---
id: 69e6284f3b7659c5a05c2a79
title: 第 5 步
challengeType: 20
dashedName: step-5
---

# --description--

你可以使用 `in` 关键字来查看字串是否存在于字串中。如果字串存在，它会返回 `True`，否则返回 `False`。下面是一个例子：

```py
sentence = 'I love Python'
print('Python' in sentence) # True
print('Java' in sentence) # False
```

创建一个名为 `has_at` 的变量，用来检查字符 `@` 是否在 `email` 中。

然后打印 `has_at`，使输出显示为：`Has @ symbol: True`。

# --hints--

你应该创建一个名为 `has_at` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_at')
`) })
```

你的 `has_at` 变量应该赋值为 `'@' in email`。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_at').is_equivalent("has_at = '@' in email")
`) })
```

你应该打印 `has_at`，使输出显示为 `Has @ symbol: True`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has @ symbol: {has_at}')",
    "print('Has @ symbol: ' + str(has_at))",
    "print('Has @ symbol:', has_at)",
    "print('Has @ symbol:', str(has_at))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

--fcc-editable-region--

--fcc-editable-region--
```
