---
id: 69e62ce96df237884da25282
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

`.lower()` 方法将字串中的所有字符转换为小写。这对于规范化电子邮件地址非常有用：

```python
upperGreet = 'HELLO WORLD'
lowerGreet = upperGreet.lower()
print(lowerGreet)  # Output: hello world
```

你可以这样对字串调用 `.strip()` 和 `.lower()`：

```python
messy = '  HELLO WORLD  '
clean = messy.strip().lower()
print(clean)  # Output: hello world
```

创建一个变量 `messy_email`，并赋值为带有两个前导空格和两个行尾空白的字串 `SARAH.DAVIS@COMPANY.COM`。

然后，对 `messy_email` 调用 `.strip().lower()`，并将结果赋值给变量 `clean_email`。

然后打印这两个变量，使输出显示为：

```md
Original email: '  SARAH.DAVIS@COMPANY.COM  '
Cleaned email: 'sarah.davis@company.com'
```

# --hints--

你应该创建一个名为 `messy_email` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_email')
`) })
```

你应该将 <code>'  SARAH.DAVIS@COMPANY.COM  '</code> 赋值给 `messy_email`。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_email').is_equivalent("messy_email = '  SARAH.DAVIS@COMPANY.COM  '")
`) })
```

你应该创建一个名为 `clean_email` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('clean_email')
`) })
```

你应该对 `messy_email` 调用 `.strip().lower()`，并将结果赋值给 `clean_email` 变量。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_email').is_equivalent("clean_email = messy_email.strip().lower()")
`) })
```

你应该打印 `messy_email`，使输出显示为 `Original email: '  SARAH.DAVIS@COMPANY.COM  '`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original email: '{messy_email}'\\")",
    "print(\\"Original email: '\\" + messy_email + \\"'\\")",
    "print(\\"Original email: '\\", messy_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

你应该打印 `clean_email`，使输出显示为 `Cleaned email: 'sarah.davis@company.com'`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned email: '{clean_email}'\\")",
    "print(\\"Cleaned email: '\\" + clean_email + \\"'\\")",
    "print(\\"Cleaned email: '\\", clean_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
