---
id: 69e736c39dfa403496e8071e
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

`.upper()` 方法将所有字符转换为大写。这对于制作名牌或标题非常有用：

```py
name = 'sarah davis'
badge = name.upper()
print(badge)  # Output: SARAH DAVIS
```

创建一个名为 `name_badge` 的变量，并将调用 `clean_name` 上的 `.upper()` 方法的结果赋值给它。

然后打印 `name_badge`，使输出显示为：`Name badge: SARAH DAVIS`。

# --hints--

你应该创建一个名为 `name_badge` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('name_badge')
`) })
```

你应该调用 `clean_name` 上的 `.upper()` 方法，并将结果赋值给 `name_badge` 变量。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('name_badge').is_equivalent("name_badge = clean_name.upper()")
`) })
```

你应该打印 `name_badge`，使输出显示为 `Name badge: SARAH DAVIS`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Name badge: {name_badge}')",
    "print('Name badge: ' + name_badge)",
    "print('Name badge:', name_badge)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

--fcc-editable-region--

--fcc-editable-region--
```
