---
id: 69e76d1997661be7f0040e2d
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

`.split(separator)` 方法会在每个 `separator` 出现的位置将字串拆分成子字串列表。例如：

```py
fruits = 'apple, banana, cherry'
parts = fruits.split(', ')
print(parts)  # Output: ['apple', 'banana', 'cherry']
```

创建一个变量 `full_address` 并赋值为字串 `123 Main Street, Springfield, IL`。然后通过以逗号加一个空格拆分，创建 `address_parts`。

然后打印 `address_parts`，使输出显示为 `Address parts: ['123 Main Street', 'Springfield', 'IL']`。

# --hints--

你应该创建一个名为 `full_address` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('full_address')
`) })
```

你应该将字串 `123 Main Street, Springfield, IL` 赋值给 `full_address`。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('full_address').is_equivalent("full_address = '123 Main Street, Springfield, IL'")
`) })
```

你应该创建一个名为 `address_parts` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('address_parts')
`) })
```

你应该对 `full_address` 使用 `.split(', ')`，并将结果赋值给 `address_parts` 变量。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('address_parts').is_equivalent("address_parts = full_address.split(', ')")
`) })
```

你应该打印 `address_parts`，使输出显示为 `Address parts: ['123 Main Street', 'Springfield', 'IL']`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Address parts: {address_parts}')",
    "print('Address parts: ' + str(address_parts))",
    "print('Address parts:', address_parts)",
    "print('Address parts:', str(address_parts))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

--fcc-editable-region--

--fcc-editable-region--
```
