---
id: 69e75dc335b6ec9903a87d3e
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

你也可以使用 `.replace()` 来移除不需要的字符。例如，从电话号码中移除短横线：

```py
phone = '456-147'
clean = phone.replace('-', '')
print(clean)  # Output: 456147
```

创建一个变量 `phone` 并赋值为字串 `123-456-7890`。然后将 `-` 替换为空字串，并将结果赋值给名为 `phone_clean` 的变量。

然后打印这两个变量，使输出显示为：

```md
Phone with dashes: 123-456-7890
Phone without dashes: 1234567890
```

# --hints--

你应该创建一个名为 `phone` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('phone')
`) })
```

你应该将字串 `123-456-7890` 赋值给 `phone`。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('phone').is_equivalent("phone = '123-456-7890'")
`) })
```

你应该创建一个名为 `phone_clean` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('phone_clean')
`) })
```

你应该对 `phone` 使用 `.replace('-', '')`，并将结果赋值给 `phone_clean` 变量。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('phone_clean').is_equivalent("phone_clean = phone.replace('-', '')")
`) })
```

你应该打印 `phone`，使输出显示为 `Phone with dashes: 123-456-7890`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Phone with dashes: {phone}')",
    "print('Phone with dashes: ' + phone)",
    "print('Phone with dashes:', phone)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

你应该打印 `phone_clean`，使输出显示为 `Phone without dashes: 1234567890`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Phone without dashes: {phone_clean}')",
    "print('Phone without dashes: ' + phone_clean)",
    "print('Phone without dashes:', phone_clean)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

--fcc-editable-region--

--fcc-editable-region--
```
