---
id: 69e73a141addc6f442924b0d
title: 第 12 步
challengeType: 20
dashedName: step-12
---

# --description--

你已经在 `email` 变量上使用过 `.find()` 和切片。现在你将对 `clean_email` 使用相同的技术。

查找 `clean_email` 中 `@` 字符的位置并保存到 `at_pos`。然后使用切片提取 `@` 之前的所有内容并保存到 `email_user`。

然后打印 `clean_email` 和 `email_user`，使输出显示为：

```md
Email: sarah.davis@company.com
Username: sarah.davis
```

# --hints--

你应该创建一个名为 `at_pos` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_pos')
`) })
```

你应该在 `clean_email` 上调用 `.find('@')` 并将结果赋值给 `at_pos` 变量。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_pos').is_equivalent("at_pos = clean_email.find('@')")
`) })
```

你应该创建一个名为 `email_user` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email_user')
`) })
```

你应该使用带有 `at_pos` 的切片操作 `clean_email`，并将结果赋值给 `email_user` 变量。

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('email_user')
assert node.is_equivalent("email_user = clean_email[:at_pos]") or node.is_equivalent("email_user = clean_email[0:at_pos]")
`) })
```

你应该打印 `clean_email`，使输出显示为 `Email: sarah.davis@company.com`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {clean_email}')",
    "print('Email: ' + clean_email)",
    "print('Email:', clean_email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

你应该打印 `email_user`，使输出显示为 `Username: sarah.davis`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {email_user}')",
    "print('Username: ' + email_user)",
    "print('Username:', email_user)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

--fcc-editable-region--

--fcc-editable-region--
```
