---
id: 69e73f4f5f986657647bbd63
title: 第 13 步
challengeType: 20
dashedName: step-13
---

# --description--

`.replace(old, new)` 方法会将字串 `old` 的每个出现位置替换为另一个字串 `new`。你也可以将它与其他方法链式调用。例如：

```py
text = 'john.doe'
name = text.replace('.', ' ').title()
print(name)  # Output: John Doe
```

将 `email_user` 中的 `.` 替换为单个空格，然后对结果调用 `.title()`。将结果赋值给名为 `display_name` 的变量。

然后打印 `display_name`，使输出显示为：`Display name: Sarah Davis`。

# --hints--

你应该创建一个名为 `display_name` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('display_name')
`) })
```

你应该对 `email_user` 使用 `.replace('.', ' ')` 和 `.title()`，并将结果赋值给 `display_name` 变量。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('display_name').is_equivalent("display_name = email_user.replace('.', ' ').title()")
`) })
```

你应该打印 `display_name`，使输出显示为 `Display name: Sarah Davis`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Display name: {display_name}')",
    "print('Display name: ' + display_name)",
    "print('Display name:', display_name)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

--fcc-editable-region--

--fcc-editable-region--
```
