---
id: 69e774b82fe0bfdb684f95b4
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

有时你需要查看一个字串是如何开始的。Python 提供了一个 `.startswith()` 方法来实现这一点。如果字串以给定的前缀开始，它会返回 `True`；否则返回 `False`。

这是一个示例：

```python
code = 'EMP-001'
print(code.startswith('EMP')) # True
print(code.startswith('001')) # False
```

创建一个变量 `employee_code` 并赋值为字串 `EMP-2024-SD`。然后创建一个变量 `starts_with_emp`，用来查看 `employee_code` 是否以 `EMP` 开头。

# --hints--

你应该创建一个名为 `employee_code` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('employee_code')
`) })
```

你应该将 `EMP-2024-SD` 赋值给 `employee_code`。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('employee_code').is_equivalent("employee_code = 'EMP-2024-SD'")
`) })
```

你应该创建一个名为 `starts_with_emp` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('starts_with_emp')
`) })
```

你应该对 `employee_code` 使用 `.startswith('EMP')`，并将结果赋值给变量 `starts_with_emp`。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('starts_with_emp').is_equivalent("starts_with_emp = employee_code.startswith('EMP')")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print(f'Address parts: {address_parts}')

rejoined = ' | '.join(address_parts)
print(f'Rejoined: {rejoined}')

--fcc-editable-region--

--fcc-editable-region--
```
