---
id: 6a71a8df10e823e77358d896
title: 第 6 步
challengeType: 20
dashedName: step-6
---

# --description--

创建另一个名为 `has_com` 的变量，用来查看 `email` 中是否包含 `.com`。

然后打印 `has_com`，使输出显示为：`Has .com: True`。

# --hints--

你应该创建一个名为 `has_com` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_com')
`) })
```

你的 `has_com` 变量的值应该是 `'.com' in email`。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_com').is_equivalent("has_com = '.com' in email")
`) })
```

你应该打印 `has_com`，使输出显示为 `Has .com: True`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has .com: {has_com}')",
    "print('Has .com: ' + str(has_com))",
    "print('Has .com:', has_com)",
    "print('Has .com:', str(has_com))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')
--fcc-editable-region--

--fcc-editable-region--
```
