---
id: 6a71b3e49668c3be461257b5
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

`.title()` 方法将每个单词的首字母大写。下面是一个例子：

```py
lowercase = 'hello world'
title_case = lowercase.title()
print(title_case)  # Output: Hello World
```

创建一个名为 `clean_name` 的变量，并将对 `messy_name` 调用 `.title()` 的结果赋值给它。

# --hints--

你应该创建一个名为 `clean_name` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('clean_name')
`) })
```

你应该对 `messy_name` 使用 `.title()`，将结果赋值给 `clean_name`。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_name').is_equivalent("clean_name = messy_name.title()")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
