---
id: 6a86b9e64da746838e67fb13
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

`.strip()` 方法会移除字串两端的空白字符。

```python
messy = '  hello world  '
clean = messy.strip()
print(clean)  # Output: hello world
```

你可以链式调用 `.strip()` 和 `.title()`，先移除字串两端的空白字符，再将每个单词的首字母大写：

```python
messy = '  hello world  '
clean = messy.strip().title()
print(clean)  # Output: Hello World
```

更新 `clean_name` 的值，使其在调用 `.title()` 之前先对 `messy_name` 调用 `.strip()`。

之后，打印 `clean_name`，使输出显示为 `Cleaned name: 'Sarah Davis'`。

# --hints--

你应该更新 `clean_name`，使其在调用 `.title()` 之前先对 `messy_name` 调用 `.strip()`。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_name').is_equivalent("clean_name = messy_name.strip().title()")
`) })
```

你应该打印 `clean_name`，使输出显示为 `Cleaned name: 'Sarah Davis'`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned name: '{clean_name}'\\")",
    "print(\\"Cleaned name: '\\" + clean_name + \\"'\\")",
    "print(\\"Cleaned name: '\\", clean_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--
clean_name = messy_name.title()
--fcc-editable-region--
```
