---
id: 6a870cc60c23f0332a2fe788
title: 第 2 步
challengeType: 20
dashedName: step-2
---

# --description--

记得在之前的工作坊中，你可以通过在开头的引号前写一个 `f`，并将变量包裹在 `{}` 中，使用 f-string 在字串中包含变量。

这里是它的工作原理的提醒：

```python
name = 'Alice'
print(f'Hello, {name}')  # Hello, Alice
```

现在，打印 `email`，使输出显示为：`Email: alice.johnson@company.com`。

# --hints--

你应该打印 `email`，使输出显示为 `Email: alice.johnson@company.com`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {email}')",
    "print('Email: ' + email)",
    "print('Email:', email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
--fcc-editable-region--

--fcc-editable-region--
```
