---
id: 66ea7e7399244c6d4b6052a1
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

워크숍의 다음 부분에서는 `Math.random()` 메서드가 어떻게 작동하는지 검토합니다.

먼저 문자열 `console.log()`를 콘솔에 기록하는 또 다른 `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."`을 추가하세요.
 
# --hints--

문자열 `console.log()`를 기록하는 `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."`이 있어야 합니다.

```js
assert.match(code, /console\.log\((`|'|")\s*The\s+Math\.random\(\)\s+method\s+returns\s+a\s+pseudo\s+random\s+number\s+greater\s+than\s+or\s+equal\s+to\s+0\s+and\s+less\s+than\s+1\.\1\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);
--fcc-editable-region--

--fcc-editable-region--
```
