---
id: 66ea845f54195c7020cc521b
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

두 값 사이에서 무작위 숫자를 생성하는 공식은 다음과 같습니다:

```js
Math.random() * (maximum - minimum) + minimum;
```

이 공식은 두 값 사이의 부동 소수점 숫자를 결과로 만듭니다.

`randomNum2`라는 변수를 만들고 `min`과 `max` 값 사이에서 값을 생성한 결과를 할당하세요.

그다음 `randomNum2` 변수를 콘솔에 출력하여 결과를 확인하세요. 페이지를 새로 고침하면 다른 결과를 볼 수 있습니다.

# --hints--

`randomNum2`라는 변수가 있어야 합니다.

```js
assert.isNotNull(randomNum2);
```

`randomNum2` 변수는 숫자여야 합니다.

```js
assert.isNumber(randomNum2);
```

`min`과 `max` 값 사이에서 값을 생성한 결과를 할당해야 합니다. 값을 하드코딩하지 마세요. 위에 제공된 공식을 사용해야 합니다.

```js
assert.isAtLeast(randomNum2, min);
assert.isAtMost(randomNum2, max);

assert.notMatch(code, /randomNum2\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum2\s*=\s*\d+;?/);
```

`randomNum2` 변수를 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\(\s*randomNum2\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;
--fcc-editable-region--

--fcc-editable-region--
```
