---
id: 66ea8239b9eeb76f12299030
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

워크숍의 다음 부분은 두 값 사이에서 무작위 숫자를 생성하는 방법을 검토하는 것입니다.

먼저, `console.log()` 메시지를 콘솔에 출력하는 또 다른 `"Now, generate a random number between two values."`를 추가하세요.

그다음, `min`라는 변수를 만들고 `1` 값을 할당하며, `max`라는 변수를 만들고 `100` 값을 할당하세요.

다음 단계에서는 이 두 값 사이에서 무작위 숫자를 생성할 것입니다.

# --hints--

`console.log()` 메시지를 출력하는 `"Now, generate a random number between two values."`가 있어야 합니다.

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+number\s+between\s+two\s+values\.\s*\1\);?/);
```

`min`라는 변수가 있어야 합니다.

```js
assert.isNotNull(min);
```

`min` 변수는 숫자여야 합니다.

```js
assert.isNumber(min);
```

`1` 변수에 숫자 `min`을 할당해야 합니다.

```js
assert.equal(min, 1);
```

`max`라는 변수가 있어야 합니다.

```js
assert.isNotNull(max);
```

`max` 변수는 숫자여야 합니다.

```js
assert.isNumber(max);
```

`100` 변수에 숫자 `max`을 할당해야 합니다.

```js
assert.equal(max, 100); 
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);
--fcc-editable-region--

--fcc-editable-region--
```
