---
id: 66ea7adae8053065a64f9002
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 Mathbot을 만들어 다양한 `Math` 객체 메서드를 다루는 방법을 검토합니다. 이 Mathbot은 JavaScript의 `Math` 객체에 대해 알려주기 위해 수학 연산과 메시지를 콘솔에 기록할 것입니다.

첫 번째 단계로, `botName`라는 변수를 만들고 문자열 값 `"MathBot"`을 할당하세요.

그다음, `greeting`라는 또 다른 변수를 만들고 문장 `"Hi there! My name is [botName goes here] and I am here to teach you about the Math object!"`를 할당하세요. `[botName goes here]` 자리 표시자 대신 `botName` 변수를 사용하세요.

템플릿 리터럴이나 `+` 연산자를 사용한 문자열 연결을 자유롭게 사용해도 됩니다.

마지막으로 `greeting` 변수를 콘솔에 기록하세요.

# --hints--

`botName`라는 변수가 있어야 합니다.

```js
assert.isNotNull(botName);
```

`botName` 변수는 문자열이어야 합니다.

```js
assert.isString(botName);
```

`"MathBot"` 변수에 문자열 값 `botName`을 할당해야 합니다.

```js
assert.strictEqual(botName, "MathBot");
```

`greeting`라는 변수가 있어야 합니다.

```js
assert.isNotNull(greeting);
```

`greeting` 변수는 문자열이어야 합니다.

```js
assert.isString(greeting);
```

다음 문장을 `greeting` 변수에 할당해야 합니다: `"Hi there! My name is [MathBot goes here] and I am here to teach you about the Math object!"`. `[MathBot goes here]`을 `botName` 변수로 바꾸는 것을 잊지 마세요.

```js
assert.strictEqual(greeting, "Hi there! My name is MathBot and I am here to teach you about the Math object!");
```

`greeting` 변수를 콘솔에 기록해야 합니다.

```js
assert.match(code, /console\.log\(greeting\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
