---
id: 66ea7fc048a6e26e33c598fe
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

이제 `Math.random()` 메서드를 사용해 무작위 숫자를 생성할 차례입니다.

`randomNum`라는 변수를 만들고 `Math.random()` 메서드를 호출한 결과 값을 할당하세요.

그다음 `randomNum` 변수를 콘솔에 출력하세요.

코드에 공백을 추가해 봇을 다시 실행하면 생성되는 무작위 숫자가 달라지는 것을 확인할 수 있습니다.

# --hints--

`randomNum`라는 변수가 있어야 합니다.

```js
assert.isNotNull(randomNum);
```

`Math.random()` 메서드를 호출한 결과를 `randomNum` 변수에 할당해야 합니다.

```js
assert.match(code, /randomNum\s*=\s*Math\.random\(\);?/);
```

`randomNum` 변수에 무작위 숫자를 하드코딩하지 마세요. 무작위 숫자를 생성하려면 `Math.random()` 메서드를 사용하세요.

```js
assert.notMatch(code, /randomNum\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum\s*=\s*\d+;?/);
```

`randomNum` 변수를 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\(\s*randomNum\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");
--fcc-editable-region--

--fcc-editable-region--
```
