---
id: 66ea89c25e1bd871688eba6a
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

지금까지는 부동 소수점 숫자와 숫자 `0`만 다뤘습니다.

워크숍의 다음 단계에서는 `Math.floor()` 메서드를 다루는 방법을 검토합니다.

먼저 `console.log()`를 추가하여 메시지 `"The Math.floor() method rounds the value down to the nearest whole integer."`를 기록하세요.

# --hints--

메시지 `"The Math.floor() method rounds the value down to the nearest whole integer."`를 기록해야 합니다.

```js
assert.match(code, /console\.log\((["'])The\s+Math\.floor\(\)\s+method\s+rounds\s+the\s+value\s+down\s+to\s+the\s+nearest\s+whole\s+integer\.\1\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);
--fcc-editable-region--

--fcc-editable-region--
```
