---
id: 66eb488c0129e192542f0865
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

이전 수업에서 `Math.ceil()` 메서드를 이렇게 사용하는 방법을 배웠습니다:

```js
const price = 10.01;
Math.ceil(price); // 11
```

`Math.ceil()` 메서드는 값을 가장 가까운 정수로 올림합니다.

`numRoundedUp` 부동 소수점 숫자를 가장 가까운 정수로 올림한 결과를 `3.2`라는 변수에 할당하세요.

그다음, 결과를 확인하기 위해 `numRoundedUp` 변수를 콘솔에 출력하세요.

# --hints--

`numRoundedUp` 변수가 있어야 합니다.

```js
assert.isNotNull(numRoundedUp);
```

`numRoundedUp` 변수는 숫자여야 합니다.

```js
assert.isNumber(numRoundedUp);
```

`numRoundedUp` 변수는 `3.2` 부동 소수점 숫자를 가장 가까운 정수로 내림한 결과와 같아야 합니다.

```js
assert.equal(numRoundedUp, 4);
```

`4` 값을 `numRoundedUp` 변수에 하드코딩하지 마세요.

```js
assert.notMatch(code, /numRoundedUp\s*=\s*4/);
```

`numRoundedUp` 변수를 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\(\s*numRoundedUp\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

--fcc-editable-region--

--fcc-editable-region--
```
