---
id: 66eb3ec3b07dfe8f898646c0
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

이전 수업에서 배운 것처럼, `Math.floor()` 메서드는 값을 가장 가까운 정수로 내림합니다.

```js
const price = 10.99;
Math.floor(price); // 10
```

`numRoundedDown`이라는 변수를 만들고, 부동 소수점 숫자 `6.7`를 가장 가까운 정수로 내림한 결과를 할당하세요.

그런 다음, 결과를 확인하기 위해 `numRoundedDown` 변수를 콘솔에 출력하세요.

# --hints--

`numRoundedDown` 변수가 있어야 합니다.

```js
assert.isNotNull(numRoundedDown);
```

`numRoundedDown` 변수는 숫자여야 합니다.

```js
assert.isNumber(numRoundedDown);
```

`numRoundedDown` 변수는 부동 소수점 숫자 `6.7`를 가장 가까운 정수로 내림한 결과와 같아야 합니다.

```js
assert.equal(numRoundedDown, 6);
```

`6` 값을 `numRoundedDown` 변수에 하드코딩하지 마세요.

```js
assert.notMatch(code, /numRoundedDown\s*=\s*6/);
```

`numRoundedDown` 변수를 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\(\s*numRoundedDown\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");
--fcc-editable-region--

--fcc-editable-region--
```
