---
id: 66ece147323c91f654bf8dd9
title: 15단계
challengeType: 1
dashedName: step-15
---

# --description--

이전 수업에서, 다음과 같이 숫자 범위에서 최댓값과 최솟값을 구하는 방법을 배웠습니다:

```js
Math.max(1, 2, 3, 4, 5); // 5
Math.min(1, 2, 3, 4, 5); // 1
```

`maxNum`라는 변수를 만들고, 숫자 `3, 125, 55, 24` 사이에서 최댓값을 찾은 결과를 할당하세요. 그런 다음 `maxNum`의 값을 콘솔에 출력하세요.

그 아래에 `minNum`라는 변수를 만들고, 숫자 `6, 90, 14, 90, 2` 사이에서 최솟값을 찾은 결과를 할당하세요. 그런 다음 `minNum`의 값을 콘솔에 출력하세요.

# --hints--

`maxNum`라는 변수가 있어야 합니다.

```js
assert.isNotNull(maxNum);
```

`maxNum` 변수에는 숫자 `3, 125, 55, 24` 사이에서 최댓값을 찾은 결과가 할당되어야 합니다.

```js
assert.equal(maxNum, 125);
```

`125` 변수에 `maxNum` 값을 하드코딩하지 마세요. 대신 `Math.max()` 메서드를 사용해야 합니다.

```js
assert.notMatch(code, /maxNum\s*=\s*125/);
```

`maxNum` 값을 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\(\s*maxNum\s*\)/);
```

`minNum`라는 변수가 있어야 합니다.

```js
assert.isNotNull(minNum);
```

`minNum` 변수에는 숫자 `6, 90, 14, 90, 2` 사이에서 최솟값을 찾은 결과가 할당되어야 합니다.

```js
assert.equal(minNum, 2);
```

`2` 변수에 `minNum` 값을 하드코딩하지 마세요. 대신 `Math.min()` 메서드를 사용해야 합니다.

```js
assert.notMatch(code, /minNum\s*=\s*2/);
```

`minNum` 값을 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\(\s*minNum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

console.log(
  "The Math.round() method rounds the value to the nearest whole integer."
);

const numRounded = Math.round(2.7);
console.log(numRounded);
const numRounded2 = Math.round(11.2);
console.log(numRounded2);

console.log("The Math.max() and Math.min() methods are used to get the maximum and minimum number from a range.");

--fcc-editable-region--

--fcc-editable-region--
```
