---
id: 69218a15fe536ee1be38aefb
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

워크숍의 다음 부분은 두 값 사이에서 무작위 정수를 생성하는 방법을 검토하는 것입니다.

먼저 `console.log()`를 하나 더 추가하여 메시지 `"Now, generate a random integer between two values."`를 콘솔에 기록하세요.

다음 단계에서는 이미 생성한 `min`와 `max` 변수 값 사이에서 무작위 정수를 생성할 것입니다.

# --hints--

메시지 `console.log()`를 콘솔에 기록하는 `"Now, generate a random integer between two values."`가 있어야 합니다.

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+integer\s+between\s+two\s+values\.\s*\1\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

--fcc-editable-region--

--fcc-editable-region--
```
