---
id: 69218e46f4490b33f2481a1a
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

두 값 사이의 무작위 정수를 생성하는 공식은 다음과 같습니다:

```js
Math.floor(Math.random() * (maximum - minimum) + minimum);
```

이 공식은 두 값 사이의 정수를 결과로 만듭니다.

`randomInt`라는 변수를 만들고 `min`과 `max` 값 사이에서 값을 생성한 결과를 할당하세요.

그런 다음 `randomInt` 변수를 콘솔에 출력하여 결과를 확인하세요. 페이지를 새로 고침하여 다른 결과를 확인해 보세요.

# --hints--

`randomInt`라는 변수가 있어야 합니다.

```js
assert.isNotNull(randomInt);
```

`randomInt` 변수는 숫자여야 합니다.

```js
assert.isNumber(randomInt);
```

`min`과 `max` 값 사이에서 값을 생성한 결과를 할당해야 합니다. 값을 하드코딩하지 마세요. 위에 제공된 공식을 사용해야 합니다.

```js
assert.isAtLeast(randomInt, min);
assert.isAtMost(randomInt, max);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomInt\s*=\s*\d+;?/);
```

`randomInt` 변수를 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\(\s*randomInt\s*\);?/);
```


# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");
--fcc-editable-region--

--fcc-editable-region--
```
