---
id: 6ac0ae563f6de70930f47ef7
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Cuando escribes código, a menudo querrás reutilizar o hacer referencia a valores existentes en tu código. Las variables te permiten hacer eso.

Para asignar un valor a una variable, escribe el nombre de la variable seguido del operador de asignación (`=`) y el valor que quieres asignar a esa variable:

```py
sentence = 'I am learning Python'
```

Ve a la línea resaltada, crea una variable llamada `sentence` y asígnale la cadena `I am learning Python`.

# --hints--

Debes tener `sentence = 'I am learning Python'` en tu código. Asegúrate de comenzar al inicio de la línea.

```js
({ test: () => runPython(`
import re
target_var = _Node(_code).find_variable("sentence")
pattern = r"""sentence\\s*=\\s*(['"])\\s*i\\s+am\\s+learning\\s+python\\s*\\1\\s*"""
assert re.fullmatch(pattern, str(target_var).strip(), flags=re.IGNORECASE)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')
--fcc-editable-region--

--fcc-editable-region--
```
