---
id: 6ac0b35f3b29a27f7e207295
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

En el paso anterior, creaste la variable `sentence`. Ahora imprime su valor en la terminal escribiendo `print(sentence)` en la línea resaltada.

Con eso, has completado tu primer programa en Python.

# --hints--

Debes tener `print(sentence)` en tu código. Asegúrate de comenzar al inicio de la línea.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(sentence)")`)) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
print(sentence)
```
