---
id: 6ac0aaba914d81b32567c076
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

`print` es una función integrada que puedes usar para enviar salida al terminal. Como puedes ver, el texto `Hello` ahora es visible en el terminal.

Ahora, imprime el texto `freeCodeCamp` en el terminal yendo a la línea resaltada y escribiendo:

```py
print('freeCodeCamp')
```

Asegúrate de comenzar al principio de la línea.

# --hints--

Debes tener `print('freeCodeCamp')` en tu código. Asegúrate de comenzar al principio de la línea.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*freecodecamp\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
--fcc-editable-region--

--fcc-editable-region--
```
