---
id: 6ac09f8420b4b67e075d2cd2
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este taller, crearás tu primer programa en Python.

Este es tu editor de código, donde escribirás código en Python. La línea 1 está resaltada. Ve allí y escribe:

```py
'Hello'
```

Asegúrate de comenzar al principio de la línea y no olvides las comillas.

Cuando termines, haz clic en el botón "check your code" para ver si es correcto.

# --hints--

Debes escribir `'Hello'`. Asegúrate de comenzar al principio de la línea y no olvides las comillas.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""(['"])\\s*hello\\s*\\1"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
