---
id: 6ac0a4463662d9738c393e21
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

En el paso anterior, creaste una **cadena**: una secuencia de caracteres encerrada en comillas simples o comillas dobles. En programación, las cadenas se usan para representar texto.

El **terminal** es donde puedes ver la salida de tu código.

Para poder ver una salida, convierte tu línea de código existente en:

```py
print('Hello')
```

Asegúrate de comenzar desde el inicio de la línea.

# --hints--

Debes tener `print('Hello')` en tu código. Asegúrate de comenzar desde el inicio de la línea.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*hello\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
'Hello'
--fcc-editable-region--
```
