---
id: 69446f18fdcb11aa516f7c5c
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將會創建一個 `countup` 函式（程式），該函式使用遞迴來產生一個從 1 到指定數字的陣列。

首先建立一個名為 `countup` 的空函式（程式），該函式接受一個單一參數。

# --hints--

你應該有一個名為 `countup` 的函式（程式）。

```js
assert.isFunction(countup);
```

你的 `countup` 函式應該接受一個單一參數。

```js
assert.strictEqual(countup.length, 1);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
