---
id: 69447054012a22f38c944ff0
title: 步驟 2
challengeType: 1
dashedName: step-2
---

# --description--

在你的 `countup` 函式中添加一個 `if` 述語。

檢查 `number` 是否小於 `1`，並傳回空的陣列。

這將作為你遞迴的基底情況。

# --hints--

當 `number` 小於 `1` 時，該函式應傳回一個空的陣列。

```js
assert.deepStrictEqual(countup(0), []);
```

當 `number` 不小於 `1` 時，該函式（程式）不應傳回任何值。

```js
assert.isUndefined(countup(1)); 
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
