---
id: 694470b9a65b278f981ec62e
title: 步驟 5
challengeType: 1
dashedName: step-5
---

# --description--

最後，透過在 `console.log()` 述語中以引數 `5` 呼叫你的 `countup` 函式來測試它。 

輸出應該是 `[1, 2, 3, 4, 5]`。

# --hints--

你應該在 `console.log()` 述語中以引數 `5` 呼叫函式（程式）`countup`。

```js
assert.deepStrictEqual(countup(5), [1, 2, 3, 4, 5]);
assert.match(code, /console\.log\s*\(\s*countup\s*\(\s*5\s*\)\s*\)/)
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
  } else {
    countArray = countup(number - 1);
    countArray.push(number);
    return countArray;
  }
}
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
  } else {
    countArray = countup(number - 1);
    countArray.push(number);
    return countArray;
  }
}
console.log(countup(5));
```
