---
id: 6948705160b66cf93a916e29
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

在 `countup` 函式（程式）內的 `if` 述語正上方宣告一個名為 `countArray` 的空陣列。

這個陣列將在每次函式呼叫時儲存 `number` 的值。

# --hints--

你應該宣告一個名為 `countArray` 的空陣列。

```js
const explorer = await __helpers.Explorer(code);
const { countup } = explorer.functions;
assert.exists(countup);
const { countArray } = countup.variables;
assert.exists(countArray);
assert.isTrue(countArray.value.matches('[]'));
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region-- 
  
--fcc-editable-region--
  if (number < 1) {
    return [];
  }
}
```
