---
id: 694870bcf3c6874112b6f951
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

對於所有大於且相等於 `1` 的數字，創建一個 `else` 區段，並在其中呼叫 `countup` 函式，使用 `number - 1`，並將結果指定給 `countArray`。 

然後使用 `.push()` 將 `number` 添加到 `countArray`，最後傳回 `countArray`。 

這確保遞迴最終會達到基底案例並終止。

# --hints--

你應該將以 `number - 1` 呼叫 `countup` 的結果指定到 `countArray`。

```js
assert.match(countup.toString(), /countArray\s*=\s*countup\(\s*number\s*-\s*1\s*\)/);
```

你應該將 `number` 推入 `countArray`。

```js
assert.match(countup.toString(), /countArray\.push\(\s*number\s*\)/);
```

你應該傳回 `countArray`。

```js
assert.deepStrictEqual(countup(3), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
