---
id: 668f08ea07b99b1f4a91acab
title: 创建一个食谱页面
challengeType: 25
dashedName: build-a-recipe-page
demoType: onClick
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个 `!DOCTYPE html` 声明。
1. 你应该有一个 `lang` 设置为 `en` 的 `html` 元素。
1. 你应该有一个 `head` 元素，其中包含一个用于你的食谱名称的 `title` 元素和一个 `charset` 属性设置为 `UTF-8` 的 `meta` 元素。
1. 你应该有一个 `body` 元素。
1. 你应该有一个包含你的食谱名称的 `h1` 元素。
1. 在 `h1` 下面，你应该有一个介绍食谱的 `p` 元素。
1. 你应该有一个文本为 `Ingredients` 的 `h2`元素，用于原料部分。
1. 在第一个 `h2` 元素下面，你应该有一个至少有四个列表项（`li` 元素）的无序列表（`ul` 元素），列出你的原料。
1. 你应该一个有文本为 `Instructions` 的第二个 `h2` 元素，用于操作步骤部分。
1. 在第二个 `h2` 元素下面，你应该有一个至少有四个列表项的有序列表（`ol` 元素），按顺序列出食谱步骤。
1. 你应该有一个 `img` 元素，它有一个设置为有效图像的 `src` 属性，如果你愿意可以使用 `https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg`，和一个描述图像的 `alt` 属性。

# --hints--

你的食谱页面该有一个 `!DOCTYPE html` 声明。

```js
assert.match(code, /<!DOCTYPE html>/i);
```

你应该有一个 `lang` 设置为 `en` 的 `html` 元素。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/gi);
```

你的 `html` 元素中应该有一个 `head` 元素。

```js
assert.match(code, /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i);
```

你的 `head` 元素中应该有一个 `title` 元素。

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i);
```

你的 `title` 元素应该包含你的食谱标题。

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

你的 `head` 元素中应该有一个 `meta` 元素。

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*meta[\s\S]*>[\s\S]*<\/\s*head\s*>/i);
```

你的 `meta` 元素应该将它的 `charset` 属性设置为 `UTF-8`。

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

你的 `html` 元素中应该有一个 `body` 元素。

```js
assert.match(code, /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i);
```

你应该有一个包含你的食谱名称的 `h1` 元素。

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

你应该只有一个 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你的 `h1` 元素下面应该有一个 `p` 元素。

```js
assert.strictEqual(document.querySelector('h1')?.nextElementSibling, document.querySelector('p'));
```

你的第一个 `p` 元素应该描述你的食谱。

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

你的第一个 `h2` 元素应该包含文本 `Ingredients`。

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Ingredients');
```

你的第一个 `h2` 元素下面应该有一个无序列表元素。

```js
assert.strictEqual(document.querySelector('ul')?.previousElementSibling.tagName, 'H2');
```

你的无序列表中应该至少有四个包含原料的列表项元素。

```js
const els = document.querySelectorAll('ul > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

你的第二个 `h2` 元素应该包含文本 `Instructions`。

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Instructions');
```

你的第二个 `h2` 元素下面应该有一个有序列表元素。

```js
assert.strictEqual(document.querySelectorAll('h2')?.[1]?.nextElementSibling?.tagName, "OL");
```

你的无序列表中应该至少有四个包含操作步骤的列表项元素。

```js
const els = document.querySelectorAll('ol > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

你应该至少有一个 `img` 元素。

```js
assert.exists(document.querySelector('img'));
```

你所有的 `img` 元素都应该有一个有效的 `src` 属性和值。

```js
const img = document.querySelector('img');
const rawSrc = img?.getAttribute('src');
const resolvedSrc = img?.src;
const re = new RegExp(window.location.href, "ig");

assert.isAbove(rawSrc?.trim().length, 0, "The 'src' attribute must be explicitly set.");
assert.notMatch(resolvedSrc, re, "The 'src' should not start with the current page URL");

img.onload = () => {
  console.log('Image loaded successfully.');
};

img.onerror = (error) => {
  console.error('Image failed to load:', error);
  assert.fail("Your image's URL should be valid."); // Make the test instafail
};

if (img.complete) {
  img.onload && img.onload();
};
```

你所有的 `img` 元素都应该有一个 `alt` 属性来描述图像。

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Chocolate chip cookies recipe</title>
</head>

<body>
  <h1>Chocolate Chip Cookies</h1>
  <p>Welcome to the ultimate guide for making mini chocolate chip cookies! These bite-sized treats are perfect for
    satisfying your sweet tooth without overindulging. Follow this simple recipe to create delicious,
    crispy-on-the-outside, chewy-on-the-inside mini chocolate chip cookies that everyone will love.</p>
  <img src="https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg" alt="Ingredients for baking: three eggs, a bowl of flour, a glass of milk, and a whisk arranged on a wooden table.">
  <h2>Ingredients</h2>
  <ul>
    <li>1 cup all-purpose flour</li>
    <li>1/2 teaspoon baking soda</li>
    <li>1/4 cup unsalted butter, softened</li>
    <li>1/4 cup granulated sugar</li>
    <li>1/2 teaspoon vanilla extract</li>
    <li>1/2 cup mini chocolate chips</li>
  </ul>
  <h2>Instructions</h2>
  <ol>
    <li>Preheat your oven to 350°F (175°C) and line a baking sheet with parchment paper.</li>
    <li>In a bowl, whisk together the flour and baking soda.</li>
    <li>In another bowl, beat the butter, sugar, and vanilla extract until creamy.</li>
    <li>Gradually add the dry ingredients to the wet mixture, then fold in the mini chocolate chips.</li>
    <li>Drop small spoonfuls of dough onto the baking sheet.</li>
    <li>Bake for 8-10 minutes, then let cool before enjoying!</li>
  </ol>
</body>

</html>
```
