---
id: 616968c2c94c8071b349c146
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

В межах `body` вкладіть елемент `h1` з текстом `CSS Color Markers`.

# --hints--

Ваш код має містити елемент `h1`.

```js
const title = document.querySelector('h1');
assert.exists(title);
```

У вас має бути відкриваючий тег `h1`.

```js
assert.match(code, /<h1\s*>/i);
```

У вас має бути закриваючий тег `h1`.

```js
assert.match(code, /<\/h1\s*>/i);
```

Ваш елемент `h1` має бути всередині елемента `body`.

```js
assert.strictEqual(document.querySelector('h1')?.parentElement?.localName, 'body');
```

Ваш елемент `h1` має містити текст `CSS Color Markers`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

Пам’ятайте, що регістр і правопис мають значення для тексту `h1`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
