---
id: 61696ef7ac756c829f9e4048
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

У цьому проєкті ви працюватимете з зовнішнім CSS-файлом для стилізації сторінки. Файл `styles.css` уже створено для вас. Але перед тим, як ви зможете його використати, потрібно зв’язати його зі сторінкою.

Вкладіть елемент `link` всередину елемента `head`. Додайте йому атрибут `rel` зі значенням `"stylesheet"` та атрибут `href` зі значенням `"styles.css"`.

# --hints--

У вас має бути один елемент `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Ваш елемент `link` має бути порожнім елементом, він не повинен мати кінцевого тегу `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Ваш елемент `link` має мати атрибут `rel` зі значенням `"stylesheet"`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

Ваш елемент `link` має мати атрибут `href` зі значенням `"styles.css"`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

Ваш елемент `link` має бути всередині елемента `head`.

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
