---
id: 616d525007b8c5d8b3308b1c
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Ваш маркер виглядатиме краще, якщо він буде по центру сторінки. Легкий спосіб зробити це — використати `margin` <dfn>скорочену властивість</dfn>.

Ви можете окремо задавати область поля елементів за допомогою властивостей, таких як `margin-top` і `margin-left`, а скорочена властивість `margin` полегшує одночасне встановлення кількох областей поля.

Щоб центровано розмістити маркер на сторінці, встановіть його властивість `margin` у значення `auto`. Це встановить `margin-top`, `margin-right`, `margin-bottom` і `margin-left` усі в `auto`.

# --hints--

Ваш CSS-правило `.marker` має містити властивість `margin`, встановлену в `auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
