---
id: 616d4a84b756d9c4b8255093
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Настав час додати кольору маркеру. Пам’ятайте, що один зі способів додати колір елементу — це використати <dfn>ключове слово кольору</dfn>, наприклад `black`, `cyan` або `yellow`.

Нагадаємо, як націлитися на клас `freecodecamp`:

```css
.freecodecamp {
  
}
```

Створіть нове CSS-правило, яке націлюється на клас `marker`, і встановіть його властивість `background-color` у значення `red`.

**Примітка:** Після додавання CSS ви не побачите жодних змін.

# --hints--

Вам слід створити селектор класу, щоб націлитися на клас `marker`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.marker'));
```

Ваше CSS-правило `.marker` має містити властивість `background-color`, встановлену в `red`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor, 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
