---
id: 616d595270d902f0e7086e18
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Хоч у вас є три окремі маркери `div` елементи, вони виглядають як один великий прямокутник. Слід додати трохи пробілу між ними, щоб було легше бачити кожен елемент.

Коли скорочена властивість `margin` має два значення, вона встановлює `margin-top` і `margin-bottom` на перше значення, а `margin-left` і `margin-right` — на друге.

У вашому правилі CSS `.marker` встановіть властивість `margin` у `10px auto`.

# --hints--

Ваше правило CSS `.marker` має містити властивість `margin`, встановлену у `10px auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

```
