---
id: 616d50b93ba424d6282c99cf
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Фоновий колір застосовано, але оскільки елемент маркера `div` не має вмісту, він за замовчуванням не має висоти.

У вашому правилі CSS `.marker` встановіть властивість `height` у значення `25px` та властивість `width` у значення `200px`.

# --hints--

Ваше правило CSS `.marker` має містити властивість `width`, встановлену у `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.width, '200px');
```

Ваше правило CSS `.marker` має містити властивість `height`, встановлену у `25px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.height, '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  background-color: red;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
