---
id: 61764c602bee6974e7790f35
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Щоб надати маркерам різні кольори, потрібно додати унікальний клас до кожного з них. Ви можете додати кілька класів до елемента, перерахувавши їх у атрибуті `class` і розділивши пробілом. Наприклад, наступний код додає класи `animal` і `dog` до елемента `div`:

```html
<div class="animal dog">
```

Якщо ви додаєте кілька класів до HTML-елемента, стилі перших класів у таблиці стилів можуть бути перекриті пізнішими класами.

Для початку додайте клас `one` до першого маркера `div` елемента.


# --hints--

Ви повинні додати клас `one` до першого маркера `div` елемента.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

Ваш перший маркер `div` має мати класи `marker` і `one`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('marker'));
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
--fcc-editable-region--
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
