---
id: 6176519636a76b810ab1219a
title: Крок 15
challengeType: 0
dashedName: step-15
---

# --description--

Створіть CSS-правило, яке націлюється на клас `two` і встановіть його властивість `background-color` у значення `green`.

Також створіть окреме CSS-правило, яке націлюється на клас `three` і встановіть його `background-color` у значення `blue`.

# --hints--

Ви повинні використати селектор класу, щоб націлитися на клас `two`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.two'));
```

Ваше CSS-правило `.two` має мати властивість `background-color`, встановлену в `green`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'green');
```

Ви повинні використати селектор класу, щоб націлитися на клас `three`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.three'));
```

Ваше CSS-правило `.three` має мати властивість `background-color`, встановлену в `blue`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

--fcc-editable-region--

--fcc-editable-region--

```
