---
id: 617bbb6b97a83f6d1f7d6e38
title: Крок 26
challengeType: 0
dashedName: step-26
---

# --description--

Тепер, коли ви попрактикувалися з вторинними кольорами, давайте повторимо, як створювати <dfn>третинні кольори</dfn>. Третинні кольори створюються шляхом поєднання первинного кольору з найближчим вторинним кольором.

Щоб створити третинний помаранчевий колір, оновіть функцію `rgb` у правилі CSS `.one`, встановивши червоний колір на максимальне значення (`255`) і зелений на `127`.

# --hints--

Ваше правило CSS `.one` має містити властивість `background-color`, встановлену в `rgb(255, 127, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
--fcc-editable-region--
  background-color: rgb(255, 255, 0);
--fcc-editable-region--
}

.two {
  background-color: rgb(0, 255, 255);
}

.three {
  background-color: rgb(255, 0, 255);
}

```
