---
id: 617bb9fdef27bc6aa0470ac2
title: Крок 25
challengeType: 0
dashedName: step-25
---

# --description--

Щоб створити фінальний вторинний колір, мадженту, оновіть функцію `rgb` у правилі CSS `.three`, щоб поєднати чистий синій та чистий червоний.

# --hints--

Ваше правило CSS `.three` має містити властивість `background-color`, встановлену в `rgb(255, 0, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(255, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 255, 0);
}

.two {
  background-color: rgb(0, 255, 255);
}

.three {
--fcc-editable-region--
  background-color: rgb(0, 0, 255);
--fcc-editable-region--
}

```
