---
id: 617b8b38f32bf2080a140675
title: Крок 17
challengeType: 0
dashedName: step-17
---

# --description--

Як ви дізналися в попередніх уроках, функція — це частина коду, яка може приймати вхідні дані та виконувати певну дію. CSS-функція `rgb` приймає значення, або <dfn>аргументи</dfn>, для червоного, зеленого та синього, і створює колір:

```css
rgb(red, green, blue);
```

Кожне значення червоного, зеленого та синього — це число від `0` до `255`. `0` означає, що цього кольору 0%, і це чорний. `255` означає, що цього кольору 100%.

У CSS-правилі `.one` замініть ключове слово кольору `red` на функцію `rgb`. Для функції `rgb` встановіть значення для червоного `255`, для зеленого — `0`, а для синього — `0`.

# --hints--

Ваше CSS-правило `.one` не повинно використовувати ключове слово кольору `red` для встановлення властивості `background-color`.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'red');
```

Ваше CSS-правило `.one` має містити властивість `background-color`, встановлену в `rgb(255, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
--fcc-editable-region--
  background-color: red;
--fcc-editable-region--
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
