---
id: 617b65579ce424bf5f02ca73
title: Крок 16
challengeType: 0
dashedName: step-16
---

# --description--

Як ви дізналися в уроках, існують дві основні моделі кольору: <dfn>адитивна</dfn> RGB (червоний, зелений, синій), яка використовується в електронних пристроях, та <dfn>субтрактивна</dfn> CMYK (бірюзовий, пурпуровий, жовтий, чорний), яка застосовується у друку.

У цьому проєкті ви працюватимете з моделлю RGB. Це означає, що кольори починаються з чорного і змінюються, коли додаються різні рівні червоного, зеленого та синього. Легкий спосіб побачити це — за допомогою CSS-функції `rgb`.

Створіть нове CSS-правило, яке націлене на клас `container` і встановіть його `background-color` у чорний колір за допомогою `rgb(0, 0, 0)`.

# --hints--

Ви повинні використати селектор класу, щоб націлитися на клас `container`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.container'));
```

Ваше CSS-правило `.container` має містити властивість `background-color`, встановлену в `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
