---
id: 618a132676346ac9f7fd59dd
title: Крок 34
challengeType: 0
dashedName: step-34
---

# --description--

Помітили, як ваші очі природно притягує червоний колір у центрі? При створенні сайту ви можете використати цей ефект, щоб привернути увагу до важливих заголовків, кнопок або посилань.

Існує кілька інших важливих поєднань кольорів, окрім доповнюючих кольорів, але ви дізнаєтеся про них трохи пізніше.

Поки що використайте функцію `rgb` у правилі CSS `.two`, щоб встановити властивість `background-color` у чорний колір.

# --hints--

Ваше правило CSS `.two` має містити властивість `background-color`, встановлену в `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  background-color: rgb(0, 255, 255);
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(0, 0, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(255, 0, 0);
--fcc-editable-region--
}

.three {
  background-color: rgb(0, 0, 0);
}

```
