---
id: 618a0927005553b74bfa05ae
title: Крок 31
challengeType: 0
dashedName: step-31
---

# --description--

Колірне коло — це коло, де схожі кольори, або <dfn>відтінки</dfn>, розташовані близько один до одного, а різні — далі. Наприклад, чистий червоний знаходиться між відтінками рожевого та помаранчевого.

Два кольори, які розташовані навпроти один одного на колірному колі, називаються <dfn>доповнюючими кольорами</dfn>. Якщо поєднати два доповнюючі кольори, вони утворюють сірий. Але коли їх розміщують поруч, ці кольори створюють сильний візуальний контраст і здаються яскравішими.

У функції `rgb` для CSS-правила `.one` встановіть значення червоного кольору на максимум `255`, щоб отримати чистий червоний. У функції `rgb` для CSS-правила `.two` встановіть значення зеленого та синього кольорів на максимум `255`, щоб отримати бірюзовий.

# --hints--

Ваше CSS-правило `.one` має містити властивість `background-color`, встановлену в `rgb(255, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 0, 0)');
```

Ваше CSS-правило `.two` має містити властивість `background-color`, встановлену в `rgb(0, 255, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
