---
id: 618a00ed1ca871a2b3aca0cb
title: Крок 30
challengeType: 0
dashedName: step-30
---

# --description--

Тепер, коли ви ознайомилися з усіма основними, вторинними та третинними кольорами на колірному колі, вам буде легше зрозуміти інші концепції теорії кольору та як вони впливають на дизайн.

Спочатку в CSS-правилах `.one`, `.two` і `.three` відрегулюйте значення у функції `rgb`, щоб властивість `background-color` кожного елемента була встановлена в чистий чорний колір. Пам’ятайте, що функція `rgb` використовує адитивну модель кольору, де кольори починаються з чорного і змінюються зі збільшенням значень червоного, зеленого та синього.

# --hints--

Ваше CSS-правило `.one` має мати властивість `background-color`, встановлену в `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(0, 0, 0)');
```

Ваше CSS-правило `.two` має мати властивість `background-color`, встановлену в `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 0, 0)');
```

Ваше CSS-правило `.three` має мати властивість `background-color`, встановлену в `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(127, 255, 0);
}

.two {
  background-color: rgb(0, 127, 255);
}

.three {
  background-color: rgb(255, 0, 127);
}
--fcc-editable-region--

```
