---
id: 61b0936d9e7fc020b43b1bb8
title: Крок 73
challengeType: 0
dashedName: step-73
---

# --description--

Здається, що ваш рукав зник, але не хвилюйтеся — він все ще тут. Сталося так, що ваша нова кришка `div` займає всю ширину маркера і штовхає рукав вниз на наступний рядок.

Це тому, що властивість `display` за замовчуванням для елементів `div` — це `block`. Тож коли два елементи `block` стоять поруч, вони накладаються як справжні блоки. Наприклад, ваші елементи маркера всі накладені один на одного.

Щоб розмістити два елементи `div` в одному рядку, встановіть їхні властивості `display` у `inline-block`.

Створіть нове правило, щоб націлити одночасно класи `cap` і `sleeve`, і встановіть `display` у `inline-block`.

# --hints--

Ви повинні використати селектор класу, щоб націлити одночасно класи `cap` і `sleeve`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve') || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap'));
```

Ваше CSS-правило має містити властивість `display`, встановлену у `inline-block`.

```js
assert.isTrue(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve')?.display === 'inline-block' || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap')?.display === 'inline-block');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
}

--fcc-editable-region--

--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
