---
id: 587d7b87367417b2b2512b40
title: 比較 `var` 和 `let` 關鍵字的範圍
challengeType: 1
forumTopicId: 301195
dashedName: compare-scopes-of-the-var-and-let-keywords
---

# --description--

如果你不熟悉 `let`，請查看<a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/explore-differences-between-the-var-and-let-keywords" target="_blank" rel="noopener noreferrer nofollow">這個關於 <code>let</code> 和 <code>var</code> 差異的挑戰</a>。

當你使用 `var` 關鍵字宣告變數時，如果在函式內宣告，則為區域的，否則為全域的。

`let` 關鍵字的行為類似，但具有一些額外的特性。當你在區塊、述語或表達式內使用 `let` 關鍵字宣告變數時，其作用域僅限於該區塊、述語或表達式。

例如：

```js
var numArray = [];
for (var i = 0; i < 3; i++) {
  numArray.push(i);
}
console.log(numArray);
console.log(i);
```

這裡主控台會顯示值 `[0, 1, 2]` 和 `3`。

使用 `var` 關鍵字，`i` 被宣告為全域的。所以當執行 `i++` 時，它會更新全域變數。這段程式碼類似於以下內容：

```js
var numArray = [];
var i;
for (i = 0; i < 3; i++) {
  numArray.push(i);
}
console.log(numArray);
console.log(i);
```

這裡主控台會顯示值 `[0, 1, 2]` 和 `3`。

如果你在使用 `i` 變數的 `for` 迴圈中創建函式並將其儲存以供稍後使用，這種行為將會導致問題。這是因為儲存的函式將永遠取用更新後的全域 `i` 變數的值。

```js
var printNumTwo;
for (var i = 0; i < 3; i++) {
  if (i === 2) {
    printNumTwo = function() {
      return i;
    };
  }
}
console.log(printNumTwo());
```

這裡主控台會顯示值 `3`。

如你所見，`printNumTwo()` 列印的是 3 而不是 2。這是因為指定給 `i` 的值已被更新，而 `printNumTwo()` 傳回的是全域的 `i`，而不是該函式在 for 迴圈中創建時 `i` 所擁有的值。`let` 關鍵字不會遵循此行為：

```js
let printNumTwo;
for (let i = 0; i < 3; i++) {
  if (i === 2) {
    printNumTwo = function() {
      return i;
    };
  }
}
console.log(printNumTwo());
console.log(i);
```

這裡主控台會顯示值 `2`，以及錯誤訊息 `i is not defined`。

`i` 未定義，因為它沒有在全域的範圍中宣告。它只在 `for` 迴圈述語中宣告。`printNumTwo()` 傳回了正確的值，因為在迴圈述語中由 `let` 關鍵字為三個不同的 `i` 變數（值分別為 0、1 和 2）創建了唯一的值。

# --instructions--

修正程式碼，使在 `if` 述語中宣告的 `i` 與函式第一行宣告的 `i` 為不同的變數。務必確保在你的程式碼中任何地方都不要使用 `var` 關鍵字。

這個練習旨在說明 `var` 和 `let` 關鍵字如何指定宣告變數的作用域之間的差異。當你撰寫類似本練習中使用的函式時，通常最好使用不同的變數名稱以避免混淆。

# --hints--

`var` 不應該存在於程式碼中。

```js
assert(!__helpers.removeJSComments(code).match(/var/g));
```

在 `if` 述語中宣告的變數 `i` 應該與字串 `block scope` 相等。

```js
assert(__helpers.removeJSComments(code).match(/(i\s*=\s*).*\s*.*\s*.*\1('|")block\s*scope\2/g));
```

`checkScope()` 應該傳回字串 `function scope`

```js
assert(checkScope() === 'function scope');
```

# --seed--

## --seed-contents--

```js
function checkScope() {
  var i = 'function scope';
  if (true) {
    i = 'block scope';
    console.log('Block scope i is: ', i);
  }
  console.log('Function scope i is: ', i);
  return i;
}
```

# --solutions--

```js
function checkScope() {
  let i = 'function scope';
  if (true) {
    let i = 'block scope';
    console.log('Block scope i is: ', i);
  }
 
  console.log('Function scope i is: ', i);
  return i;
}
```
