---
id: 587d7b88367417b2b2512b44
title: 撰寫帶有參數的箭頭函式（Arrow Functions）
challengeType: 1
forumTopicId: 301223
dashedName: write-arrow-functions-with-parameters
---

# --description--

就像一般函式（程式）一樣，你可以將引數傳入箭頭函式（程式）。

```js
const doubler = (item) => item * 2;
doubler(4);
```

`doubler(4)` 會傳回值 `8`。

如果箭頭函式只有一個參數，包住該參數的小括弧可以省略。

```js
const doubler = item => item * 2;
```

可以將多個引數傳遞到箭頭函式中。

```js
const multiplier = (item, multi) => item * multi;
multiplier(4, 2);
```

`multiplier(4, 2)` 會傳回值 `8`。

# --instructions--

重寫會將 `arr2` 內容附加到 `arr1` 的 `myConcat` 函式，使該函式使用箭頭函式語法。

# --hints--

你應該取代 `var` 關鍵字。

```js
assert.notMatch(code, /var/g);
```

`myConcat` 應該是一個常數變數（使用 `const`）。

```js
assert.match(code, /const\s+myConcat/g);
```

`myConcat` 應該是一個具有兩個參數的箭頭函式。

```js
assert(
  /myConcat=\(\w+,\w+\)=>/.test(__helpers.removeJSComments(code).replace(/\s/g, '')) &&
    typeof myConcat === 'function'
);
```

`myConcat()` 應該傳回 `[1, 2, 3, 4, 5]`。

```js
assert.deepEqual(myConcat([1, 2], [3, 4, 5]), [1, 2, 3, 4, 5]);
```

不應該使用 `function` 關鍵字。

```js
assert.notMatch(code, /function/g);
```

# --seed--

## --seed-contents--

```js
var myConcat = function(arr1, arr2) {
  return arr1.concat(arr2);
};

console.log(myConcat([1, 2], [3, 4, 5]));
```

# --solutions--

```js
const myConcat = (arr1, arr2) =>  {
  return arr1.concat(arr2);
};

console.log(myConcat([1, 2], [3, 4, 5]));
```
