---
id: 587d7b8d367417b2b2512b59
title: 匯入預設匯出
challengeType: 1
forumTopicId: 301205
dashedName: import-a-default-export
---

# --description--

在上一個挑戰中，你學習了 `export default` 及其用途。要匯入預設匯出，你需要使用不同的 `import` 語法。在以下範例中，`add` 是 `math_functions.js` 檔案的預設匯出。以下是如何匯入它：

```js
import add from "./math_functions.js";
```

語法在一個關鍵地方有所不同。匯入的值 `add` 並未被大括號（`{}`）包圍。這裡的 `add` 只是 `math_functions.js` 檔案預設匯出的任何內容的變數名稱。當匯入預設值時，你可以在這裡使用任何名稱。

# --instructions--

在以下程式碼中，從與此檔案相同目錄中的 `math_functions.js` 檔案匯入預設匯出。將匯入命名為 `subtract`。

# --hints--

你應該正確地從 `math_functions.js` 匯入 `subtract`。

```js
assert(__helpers.removeJSComments(code).match(/import\s+subtract\s+from\s+('|")\.\/math_functions\.js\1/g));
```

# --seed--

## --seed-contents--

```js
  
// Only change code above this line

subtract(7,4);
```

# --solutions--

```js
import subtract from "./math_functions.js";

subtract(7,4);
```
