---
id: 587d7b8b367417b2b2512b50
title: 使用 ES6 撰寫簡潔的宣告式函式（程式）
challengeType: 1
forumTopicId: 301224
dashedName: write-concise-declarative-functions-with-es6
---

# --description--

在 ES5 中於物件內定義函式（程式）時，我們必須使用關鍵字 `function`，如下所示：

```js
const person = {
  name: "Taylor",
  sayHello: function() {
    return `Hello! My name is ${this.name}.`;
  }
};
```

在 ES6 中，當你在物件中定義函式（程式）時，可以完全移除 `function` 關鍵字和冒號。以下是這種語法的範例：

```js
const person = {
  name: "Taylor",
  sayHello() {
    return `Hello! My name is ${this.name}.`;
  }
};
```

# --instructions--

將物件 `bicycle` 內的函式 `setGear` 重構為使用上述簡寫語法。

# --hints--

不應該使用傳統函式（程式）表達式。

```js
assert(!__helpers.removeJSComments(code).match(/function/));
```

`setGear` 應該是一個宣告式函式（程式）。

```js
assert(
  typeof bicycle.setGear === 'function' && __helpers.removeJSComments(code).match(/setGear\s*\(.+\)\s*\{/)
);
```

`bicycle.setGear(48)` 應該將 `gear` 值更改為 48。

```js
bicycle.setGear(48);
assert(bicycle.gear === 48);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const bicycle = {
  gear: 2,
  setGear: function(newGear) {
    this.gear = newGear;
  }
};
// Only change code above this line
bicycle.setGear(3);
console.log(bicycle.gear);
```

# --solutions--

```js
const bicycle = {
  gear: 2,
  // setGear: function(newGear) {
  setGear(newGear) {
    this.gear = newGear;
  }
};
bicycle.setGear(3);
```
