---
id: 587d7b88367417b2b2512b46
title: 為你的函式（程式）設定預設參數
challengeType: 1
forumTopicId: 301209
dashedName: set-default-parameters-for-your-functions
---

# --description--

為了幫助我們創建更靈活的函式（程式），ES6 為函式（程式）介紹了 <dfn>預設參數</dfn>。

看看這段程式碼：

```js
const greeting = (name = "Anonymous") => "Hello " + name;

console.log(greeting("John"));
console.log(greeting());
```

主控台將顯示字串 `Hello John` 和 `Hello Anonymous`。

當引數未指定（即未定義）時，預設參數會生效。如你在上述範例中所見，當你未提供參數 `name` 的值時，該參數會接收其預設值 `Anonymous`。你可以為任意多個參數添加預設值。

# --instructions--

修改函式（程式）`increment`，為其添加預設參數，使其在未指定 `value` 時會將 `number` 加 1。

# --hints--

`increment(5, 2)` 的結果應該是 `7`。

```js
assert(increment(5, 2) === 7);
```

`increment(5)` 的結果應該是 `6`。

```js
assert(increment(5) === 6);
```

`value` 的預設參數值應為 `1`。

```js
assert(__helpers.removeJSComments(code).match(/value\s*=\s*1/g));
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const increment = (number, value) => number + value;
// Only change code above this line
```

# --solutions--

```js
const increment = (number, value = 1) => number + value;
```
