---
id: 587d7b8c367417b2b2512b57
title: 使用 * 為你的檔案匯入所有內容
challengeType: 1
forumTopicId: 301210
dashedName: use--to-import-everything-from-a-file
---

# --description--

假設你有一個檔案並且你希望將它的所有內容匯入到目前的檔案中。這可以使用 `import * as` 語法來完成。以下是一個範例，示範如何將同一目錄中名為 `math_functions.js` 的檔案內容匯入到另一個檔案中：

```js
import * as myMathModule from "./math_functions.js";
```

上述的 `import` 述語會創建一個名為 `myMathModule` 的物件。這只是變數名稱，你可以命名為任何名稱。該物件將包含來自 `math_functions.js` 的所有匯出，因此你可以像存取其他物件屬性一樣存取這些函式（程式）。以下是你如何使用被匯入的 `add` 和 `subtract` 函式（程式）：

```js
myMathModule.add(2,3);
myMathModule.subtract(5,3);
```

# --instructions--

此檔案中的程式碼需要位於與目前檔案相同目錄中的檔案 `string_functions.js` 的內容。使用 `import * as` 語法將該檔案中的所有內容匯入到名為 `stringFunctions` 的物件中。

# --hints--

你的程式碼應該正確使用 `import * as` 語法。

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*\*\s*as\s+stringFunctions\s+from\s*('|")\.\/string_functions\.js\1/g
  )
);
```

# --seed--

## --seed-contents--

```js

// Only change code above this line

stringFunctions.uppercaseString("hello");
stringFunctions.lowercaseString("WORLD!");
```

# --solutions--

```js
import * as stringFunctions from "./string_functions.js";

// add code above this line
stringFunctions.uppercaseString("hello");
stringFunctions.lowercaseString("WORLD!");
```
