---
id: 587d7b89367417b2b2512b49
title: 使用解構指派從物件指定變數
challengeType: 1
forumTopicId: 301215
dashedName: use-destructuring-assignment-to-assign-variables-from-objects
---

# --description--

解構賦值允許你在擷取值時指定新的變數名稱。你可以在指定值時，在冒號後面放入新的名稱來做到這一點。

使用上一個範例中的相同物件：

```js
const user = { name: 'John Doe', age: 34 };
```

以下是你如何在指派中給予新的變數名稱的方法：

```js
const { name: userName, age: userAge } = user;
```

你可以將它讀作「取得 `user.name` 的值並指定給一個名為 `userName` 的新變數」等等。`userName` 的值會是字串 `John Doe`，而 `userAge` 的值會是數字 `34`。

# --instructions--

將兩個指派取代為等同的解構指派。它仍然應該將變數 `highToday` 和 `highTomorrow` 指定為 `HIGH_TEMPERATURES` 物件中 `today` 和 `tomorrow` 的值。

# --hints--

你應該移除 ES5 指派語法。

```js
assert(
  !__helpers.removeJSComments(code).match(/highToday = HIGH_TEMPERATURES\.today/g) &&
    !__helpers.removeJSComments(code).match(/highTomorrow = HIGH_TEMPERATURES\.tomorrow/g)
);
```

你應該使用解構賦值來創建 `highToday` 變數。

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*(today\s*:\s*highToday[^}]*|[^,]*,\s*today\s*:\s*highToday\s*)}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g
  )
);
```

你應該使用解構賦值來創建 `highTomorrow` 變數。

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*(tomorrow\s*:\s*highTomorrow[^}]*|[^,]*,\s*tomorrow\s*:\s*highTomorrow\s*)}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g
  )
);
```

`highToday` 應該與 `77` 相等，且 `highTomorrow` 應該與 `80` 相等。

```js
assert(highToday === 77 && highTomorrow === 80);
```

# --seed--

## --seed-contents--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

// Only change code below this line
  
const highToday = HIGH_TEMPERATURES.today;
const highTomorrow = HIGH_TEMPERATURES.tomorrow; 

// Only change code above this line
```

# --solutions--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

const { today: highToday, tomorrow: highTomorrow } = HIGH_TEMPERATURES;
```
