---
id: 587d7b8a367417b2b2512b4e
title: 使用模板字面值創建字串
challengeType: 1
forumTopicId: 301200
dashedName: create-strings-using-template-literals
---

# --description--

ES6 的一個新特性是 <dfn>模板字面值</dfn>。這是一種特殊的字串型別，使建立複雜字串變得更容易。

模板字面值允許你創建多行字串並使用字串插值功能來創建字串。

請參考以下程式碼：

```js
const person = {
  name: "Zodiac Hasbro",
  age: 56
};

const greeting = `Hello, my name is ${person.name}!
I am ${person.age} years old.`;

console.log(greeting);
```

主控台將顯示字串 `Hello, my name is Zodiac Hasbro!` 和 `I am 56 years old.`。

那裡發生了很多事情。首先，範例使用反引號（```），而不是引號（`'` 或 `"`）來包裹字串。其次，注意字串在程式碼和輸出中都是多行的。這樣就省去了在字串中插入 `\n`。上面使用的 `${variable}` 語法是一個佔位符。基本上，你不再需要使用 `+` 運算子來串接字串。為了將變數添加到字串中，你只需將變數放入模板字串，並用 `${` 和 `}` 包裹。同樣地，你也可以在字串字面值中包含其他表達式，例如 `${a + b}`。這種創建字串的新方式，讓你有更多彈性來創建強固的字串。

# --instructions--

使用帶有反引號的模板字面語法來創建一個包含列表元素（`li`）字串的陣列。每個列表元素的文字應該是來自 `result` 物件中 `failure` 屬性的陣列元素之一，並且具有值為 `text-warning` 的 `class` 屬性。`makeList` 函式（程式）應該傳回該列表項目字串的陣列。

使用迭代器方法（任何種類的迴圈）來取得所需的輸出（如下所示）。

```js
[
  '<li class="text-warning">no-var</li>',
  '<li class="text-warning">var-on-top</li>',
  '<li class="text-warning">linebreak</li>'
]
```

# --hints--

`failuresList` 應該是一個包含 `result failure` 訊息的陣列。

```js
assert(
  typeof makeList(result.failure) === 'object' && failuresList.length === 3
);
```

`failuresList` 應該與指定的輸出相等。

```js
assert(
  makeList(result.failure).every(
    (v, i) =>
      v === `<li class="text-warning">${result.failure[i]}</li>` ||
      v === `<li class='text-warning'>${result.failure[i]}</li>`
  )
);
```

應該使用模板字串和表達式插值。

```js
assert.match(code, /(`.*\${.*}.*`)/);
```

應該使用迭代器。

```js
assert(__helpers.removeJSComments(code).match(/for|map|reduce|forEach|while/));
```

# --seed--

## --seed-contents--

```js
const result = {
  success: ["max-length", "no-amd", "prefer-arrow-functions"],
  failure: ["no-var", "var-on-top", "linebreak"],
  skipped: ["no-extra-semi", "no-dup-keys"]
};
function makeList(arr) {
  // Only change code below this line
  const failureItems = [];
  // Only change code above this line

  return failureItems;
}

const failuresList = makeList(result.failure);
```

# --solutions--

```js
const result = {
  success: ["max-length", "no-amd", "prefer-arrow-functions"],
  failure: ["no-var", "var-on-top", "linebreak"],
  skipped: ["no-extra-semi", "no-dup-keys"]
};
function makeList(arr) {
  return arr.map(val => `<li class="text-warning">${val}</li>`);
}

const failuresList = makeList(result.failure);
```
