---
id: 587d7b8c367417b2b2512b56
title: 使用匯出來分享一個程式碼區塊
challengeType: 1
forumTopicId: 301219
dashedName: use-export-to-share-a-code-block
---

# --description--

想像有一個名為 `math_functions.js` 的檔案，裡面包含幾個與數學操作相關的函式。其中一個儲存在變數 `add` 中，該變數接受兩個數字並傳回它們的總和。你想在多個不同的 JavaScript 檔案中使用這個函式。為了與這些其他檔案共享它，你首先需要 `export` 它。

```js
export const add = (x, y) => {
  return x + y;
}
```

以上是匯出單一函式（程式）的常見方式，但你也可以這樣達成相同的效果：

```js
const add = (x, y) => {
  return x + y;
}

export { add };
```

當你匯出一個變數或函式時，你可以在另一個檔案中匯入它並使用，而不必重寫程式碼。你可以透過重複第一個範例中你想匯出的每一個項目來匯出多個東西，或者將它們全部放在第二個範例的匯出述語中，如下所示：

```js
export { add, subtract };
```

# --instructions--

編輯器中有兩個與字串相關的函式（程式）。使用你選擇的方法將它們兩個都匯出。

# --hints--

你應該正確地匯出 `uppercaseString`。

```js
assert(
  __helpers.removeJSComments(code).match(
    /(export\s+const\s+uppercaseString|export\s*{\s*(uppercaseString[^}]*|[^,]*,\s*uppercaseString\s*)})/g
  )
);
```

你應該正確地匯出 `lowercaseString`。

```js
assert(
  __helpers.removeJSComments(code).match(
    /(export\s+const\s+lowercaseString|export\s*{\s*(lowercaseString[^}]*|[^,]*,\s*lowercaseString\s*)})/g
  )
);
```

# --seed--

## --seed-contents--

```js
const uppercaseString = (string) => {
  return string.toUpperCase();
}

const lowercaseString = (string) => {
  return string.toLowerCase()
}
```

# --solutions--

```js
export const uppercaseString = (string) => {
  return string.toUpperCase();
}

export const lowercaseString = (string) => {
  return string.toLowerCase()
}
```
