---
id: 587d7b88367417b2b2512b47
title: 在函式（程式）參數中使用剩餘參數
challengeType: 1
forumTopicId: 301221
dashedName: use-the-rest-parameter-with-function-parameters
---

# --description--

為了幫助我們創建更靈活的函式，ES6 介紹了用於函式參數的 <dfn>rest parameter</dfn>。使用 rest parameter，你可以創建接受可變數量引數的函式。這些引數會被存放在一個陣列中，稍後可以從函式內部存取。

看看這段程式碼：

```js
function howMany(...args) {
  return "You have passed " + args.length + " arguments.";
}
console.log(howMany(0, 1, 2));
console.log(howMany("string", null, [1, 2, 3], { }));
```

主控台會顯示字串 `You have passed 3 arguments.` 和 `You have passed 4 arguments.`。

剩餘參數消除了使用 `arguments` 物件的需求，並允許我們對傳遞給函式 `howMany` 的參數陣列使用陣列方法。

# --instructions--

使用剩餘參數修改函式（程式）`sum`，使函式（程式）`sum`能夠接受任意數量的引數並傳回它們的總和。

# --hints--

`sum(0,1,2)` 的結果應該是 3

```js
assert(sum(0, 1, 2) === 3);
```

`sum(1,2,3,4)` 的結果應該是 10

```js
assert(sum(1, 2, 3, 4) === 10);
```

`sum(5)` 的結果應該是 5

```js
assert(sum(5) === 5);
```

`sum()` 的結果應該是 0

```js
assert(sum() === 0);
```

`sum` 應該是一個箭頭函式（arrow function），並且在 `args` 參數上使用剩餘參數語法（rest parameter syntax）（`...`）。

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/sum=\(\.\.\.args\)=>/));
```

# --seed--

## --seed-contents--

```js
const sum = (x, y, z) => {
  const args = [x, y, z];
  let total = 0;
  for (let i = 0; i < args.length; i++) {
    total += args[i];
  }
  return total;
}
```

# --solutions--

```js
const sum = (...args) => {
  let total = 0;
  for (let i = 0; i < args.length; i++) {
    total += args[i];
  }
  return total;
}
```
