---
id: 587d7b8c367417b2b2512b54
title: 使用取值函式和設值函式來控制對物件的存取
challengeType: 1
forumTopicId: 301220
dashedName: use-getters-and-setters-to-control-access-to-an-object
---

# --description--

你可以從物件取得值，並設定物件中屬性的值。

這些傳統上稱為 <dfn>取值函式</dfn> 和 <dfn>設值函式</dfn>。

取值函式的目的是簡單地將物件私有變數的值傳回給使用者，而不讓使用者直接存取該私有變數。

設值函式是用來根據傳入設值函式的值修改（設定）物件私有變數的值。這種變更可能涉及計算，甚至完全覆寫先前的值。  

```js
class Book {
  constructor(author) {
    this._author = author;
  }
  // getter
  get writer() {
    return this._author;
  }
  // setter
  set writer(updatedAuthor) {
    this._author = updatedAuthor;
  }
}
const novel = new Book('anonymous');
console.log(novel.writer);
novel.writer = 'newAuthor';
console.log(novel.writer);
```

主控台會顯示字串 `anonymous` 和 `newAuthor`。

注意用來喚起取值函式和設值函式的語法。它們甚至看起來不像函式（程式）。取值函式和設值函式很重要，因為它們隱藏了內部實作細節。

**注意：**慣例是在私有變數名稱前加上底線（`_`）。然而，這個做法本身並不會使變數成為私有。

# --instructions--

使用 `class` 關鍵字來建立 `Thermostat` 類別。`constructor` 接受華氏溫度。

在該類別中，創建一個 `getter` 以取得攝氏溫度，並創建一個接受攝氏溫度的 `setter`。

請記住 `C = 5/9 * (F - 32)` 和 `F = C * 9.0 / 5 + 32`，其中 `F` 是華氏溫度的值，`C` 是相同溫度的攝氏值。

**注意：** 當你實作這個時，你將在類別內以一種溫標追蹤溫度，無論是華氏還是攝氏。

這就是取值函式和設值函式的威力。你正在為另一個使用者建立 API，無論你追蹤哪一個，他都能取得正確的結果。

換句話說，你正在將實作細節從使用者抽象化。

# --hints--

`Thermostat` 應該是一個具有定義好的 `constructor` 方法的 `類別`。

```js
assert.isFunction(Thermostat);
assert.isFunction(Thermostat?.constructor);
```

應該使用 `class` 關鍵字。

```js
assert.match(code, /class/);
```

`Thermostat` 應該能夠被具現化。

```js
const _t = new Thermostat(122);
assert.isObject(_t);
```

當以華氏溫度值具現化 `Thermostat` 時，應該設定正確的 `temperature`。

```js
const _t = new Thermostat(122);
assert.strictEqual(_t?.temperature, 50);
```

應該定義一個 `getter`。

```js
const _desc = Object.getOwnPropertyDescriptor(Thermostat.prototype, 'temperature');
assert.isFunction(_desc?.get);
```

應該定義一個 `setter`。

```js
const _desc = Object.getOwnPropertyDescriptor(Thermostat.prototype, 'temperature');
assert.isFunction(_desc?.set);
```

使用攝氏值呼叫 `setter` 應該會設定 `temperature`。

```js
const _t = new Thermostat(32);
_t.temperature = 26;
const _u = new Thermostat(32);
_u.temperature = 50;
assert.approximately(_t.temperature, 26, 0.1);
assert.approximately(_u.temperature, 50, 0.1);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line

// Only change code above this line

const thermos = new Thermostat(76); // Setting in Fahrenheit scale
let temp = thermos.temperature; // 24.44 in Celsius
thermos.temperature = 26;
temp = thermos.temperature; // 26 in Celsius
```

# --solutions--

```js
class Thermostat {
  constructor(fahrenheit) {
    this._tempInCelsius = 5/9 * (fahrenheit - 32);
  }
  get temperature(){
    return this._tempInCelsius;
  }
  set temperature(newTemp){
    this._tempInCelsius = newTemp;
  }
}

const thermos = new Thermostat(76); // Setting in Fahrenheit scale
let temp = thermos.temperature; // 24.44 in Celsius
thermos.temperature = 26;
temp = thermos.temperature; // 26 in Celsius
```
