---
id: 587d7b8c367417b2b2512b58
title: 使用 export default 創建匯出備援
challengeType: 1
forumTopicId: 301199
dashedName: create-an-export-fallback-with-export-default
---

# --description--

在 `export` 課程中，你學習了稱為 <dfn>named export</dfn> 的語法。這讓你能將多個函式（程式）和變數提供給其他檔案使用。

還有一種你需要知道的 `export` 語法，稱為 <dfn>export default</dfn>。通常如果只有一個值從檔案中被匯出，你會使用這種語法。它也用於為檔案或模組創建一個後備值。

以下是使用 `export default` 的範例：

```js
export default function add(x, y) {
  return x + y;
}

export default function(x, y) {
  return x + y;
}
```

第一個是具名函式（程式），第二個是匿名函式（程式）。

由於 `export default` 用於宣告模組或檔案的預設值，你在每個模組或檔案中只能有一個值作為預設匯出。此外，你不能將 `export default` 與 `var`、`let` 或 `const` 一起使用。

# --instructions--

以下函式（程式）應該是該模組的備用值。請為此添加必要的程式碼。

# --hints--

你的程式碼應該使用 `export` 備援。

```js
assert(
  __helpers.removeJSComments(code).match(
    /export\s+default\s+function(\s+subtract\s*|\s*)\(\s*x,\s*y\s*\)\s*{/g
  )
);
```

# --seed--

## --seed-contents--

```js
function subtract(x, y) {
  return x - y;
}
```

# --solutions--

```js
export default function subtract(x, y) {
  return x - y;
}
```
