---
id: 587d7b8b367417b2b2512b53
title: 使用類別 Syntax 定義建構式函式（程式）
challengeType: 1
forumTopicId: 301212
dashedName: use-class-syntax-to-define-a-constructor-function
---

# --description--

ES6 提供了一種使用 <dfn>class</dfn> 關鍵字來創建物件的新語法。

在 ES5 中，可以透過定義 `constructor` 函式（程式）並使用 `new` 關鍵字來實例化物件。

在 ES6 中，`class` 宣告式具有一個使用 `new` 關鍵字喚起的 `constructor` 方法。如果沒有明確定義 `constructor` 方法，則會隱式定義一個沒有引數的 `constructor` 方法。

```js
// Explicit constructor
class SpaceShuttle {
  constructor(targetPlanet) {
    this.targetPlanet = targetPlanet;
  }
  takeOff() {
    console.log("To " + this.targetPlanet + "!");
  }
}

// Implicit constructor 
class Rocket {
  launch() {
    console.log("To the moon!");
  }
}

const zeus = new SpaceShuttle('Jupiter');
// prints To Jupiter! in console
zeus.takeOff();

const atlas = new Rocket();
// prints To the moon! in console
atlas.launch();
```

應注意，`class` 關鍵字宣告一個新的函式（程式），並為其添加一個建構式。當呼叫 `new` 以創建一個新的物件時，會喚起此建構式。

**注意：** 根據慣例，ES6 類別名稱應使用 UpperCamelCase，如上面使用的 `SpaceShuttle`。

`constructor` 方法是一個用於創建並初始化由類別建立的物件的特殊方法。你將在 JavaScript 演算法與資料結構認證的物件導向程式設計章節中進一步了解它。

# --instructions--

使用 `class` 關鍵字並撰寫 `constructor` 來建立 `Vegetable` 類別。

`Vegetable` 類別允許你創建一個具有 `name` 屬性的蔬菜物件，該屬性會傳遞給 `constructor`。

# --hints--

`Vegetable` 應該是一個具有定義好的 `constructor` 方法的 `class`。

```js
assert(
  typeof Vegetable === 'function' && typeof Vegetable.constructor === 'function'
);
```

應該使用 `class` 關鍵字。

```js
assert(__helpers.removeJSComments(code).match(/class/g));
```

`Vegetable` 應該能夠被具現化。

```js
assert(() => {
  const a = new Vegetable('apple');
  return typeof a === 'object';
});
```

`carrot.name` 應該傳回 `carrot`。

```js
assert(carrot.name == 'carrot');
```

# --seed--

## --seed-contents--

```js
// Only change code below this line

// Only change code above this line

const carrot = new Vegetable('carrot');
console.log(carrot.name); // Should display 'carrot'
```

# --solutions--

```js
class Vegetable {
  constructor(name) {
    this.name = name;
  }
}
const carrot = new Vegetable('carrot');
```
