---
id: 587d7b8c367417b2b2512b55
title: 使用 `import` 重複使用 JavaScript 程式碼
challengeType: 1
forumTopicId: 301208
dashedName: reuse-javascript-code-using-import
---

# --description--

`import` 允許你選擇要載入檔案或模組的哪些零件。在前一課中，範例從 `math_functions.js` 檔案匯出了 `add`。以下是你如何將它匯入以在另一個檔案中使用：

```js
import { add } from './math_functions.js';
```

在這裡，`import` 會在 `math_functions.js` 中找到 `add`，只匯入該函式供你使用，並忽略其他部分。`./` 告訴 `import` 在與目前檔案相同的資料夾中尋找 `math_functions.js` 檔案。使用這種方式的 `import` 時，必須包含相對檔案路徑（`./`）和副檔名（`.js`）。

你可以透過在 `import` 述語中這樣添加，從檔案匯入多於一個的項目：

```js
import { add, subtract } from './math_functions.js';
```

# --instructions--

為你的目前檔案添加適當的 `import` 述語，以便使用你在上一課匯出的 `uppercaseString` 和 `lowercaseString` 函式（程式）。這些函式（程式）位於名為 `string_functions.js` 的檔案中，該檔案與目前檔案位於相同目錄。

# --hints--

你應該正確地匯入 `uppercaseString`。

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*{\s*(uppercaseString[^}]*|[^,]*,\s*uppercaseString\s*)}\s+from\s+('|")\.\/string_functions\.js\2/g
  )
);
```

你應該正確地匯入 `lowercaseString`。

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*{\s*(lowercaseString[^}]*|[^,]*,\s*lowercaseString\s*)}\s+from\s+('|")\.\/string_functions\.js\2/g
  )
);
```

# --seed--

## --seed-contents--

```js
  
// Only change code above this line

uppercaseString("hello");
lowercaseString("WORLD!");
```

# --solutions--

```js
import { uppercaseString, lowercaseString } from './string_functions.js';

uppercaseString("hello");
lowercaseString("WORLD!");
```
