---
id: 587d7b89367417b2b2512b4b
title: 使用解構指派從陣列指定變數
challengeType: 1
forumTopicId: 301213
dashedName: use-destructuring-assignment-to-assign-variables-from-arrays
---

# --description--

ES6 讓解構陣列和解構物件一樣簡單。

展開運算子與陣列解構之間的一個主要差異是，展開運算子會將陣列的所有內容展開成以逗號分隔的列表。因此，你無法挑選或選擇要指定給變數的元素。

解構陣列讓我們正是這麼做：

```js
const [a, b] = [1, 2, 3, 4, 5, 6];
console.log(a, b);
```

主控台會顯示 `a` 和 `b` 的值為 `1, 2`。

變數 `a` 被指定為陣列的第一個值，變數 `b` 被指定為陣列的第二個值。我們也可以透過解構，使用逗號來達到想要的索引，從而存取陣列中任何索引的值：

```js
const [a, b,,, c] = [1, 2, 3, 4, 5, 6];
console.log(a, b, c);
```

主控台會顯示 `a`、`b` 和 `c` 的值為 `1, 2, 5`。

# --instructions--

使用解構指派來交換 `a` 和 `b` 的值，使 `a` 接收儲存在 `b` 中的值，`b` 接收儲存在 `a` 中的值。

# --hints--

交換後，`a` 的值應該是 `6`。

```js
assert(a === 6);
```

交換後，`b` 的值應該是 `8`。

```js
assert(b === 8);
```

你應該使用陣列解構來交換 `a` 和 `b`。

```js
assert(/\[\s*(\w)\s*,\s*(\w)\s*\]\s*=\s*\[\s*\2\s*,\s*\1\s*\]/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let a = 8, b = 6;
// Only change code below this line
```

# --solutions--

```js
let a = 8, b = 6;
[a, b] = [b, a];
```
