---
id: 587d7b8a367417b2b2512b4f
title: 使用物件屬性簡寫撰寫簡潔的物件宣告式
challengeType: 1
forumTopicId: 301225
dashedName: write-concise-object-literal-declarations-using-object-property-shorthand
---

# --description--

ES6 為輕鬆定義物件字面值添加了一些很好的支援。

請考慮以下程式碼：

```js
const getMousePosition = (x, y) => ({
  x: x,
  y: y
});
```

`getMousePosition` 是一個簡單的函式（程式），傳回一個包含兩個屬性的物件。ES6 提供了語法糖，消除了必須寫 `x: x` 的冗贅。你只需寫一次 `x`，它在底層會被轉換成 `x: x`（或等價的東西）。以下是使用這個新語法重寫的相同函式（程式）：

```js
const getMousePosition = (x, y) => ({ x, y });
```

# --instructions--

使用物件屬性簡寫與物件字面值來創建並傳回具有 `name`、`age` 和 `gender` 屬性的物件。

# --hints--

`createPerson("Zodiac Hasbro", 56, "male")` 應該傳回 `{name: "Zodiac Hasbro", age: 56, gender: "male"}`。

```js
assert.deepEqual(
  { name: 'Zodiac Hasbro', age: 56, gender: 'male' },
  createPerson('Zodiac Hasbro', 56, 'male')
);
```

你的程式碼不應該使用 `key:value`。

```js
assert(!__helpers.removeJSComments(code).match(/:/g))
```

# --seed--

## --seed-contents--

```js
const createPerson = (name, age, gender) => {
  // Only change code below this line
  return {
    name: name,
    age: age,
    gender: gender
  };
  // Only change code above this line
};
```

# --solutions--

```js
const createPerson = (name, age, gender) => {
  // Only change code below this line
  /*return {
    name: name,
    age: age,
    gender: gender
  };*/
  return {
    name,
    age,
    gender
  };
  // Only change code above this line
};
```
