---
id: 587d7b89367417b2b2512b4a
title: 使用解構指派從巢狀物件指定變數
challengeType: 1
forumTopicId: 301214
dashedName: use-destructuring-assignment-to-assign-variables-from-nested-objects
---

# --description--

你可以使用前兩課的相同原則，從巢狀物件中解構值。

使用類似先前範例的物件：

```js
const user = {
  johnDoe: { 
    age: 34,
    email: 'johnDoe@freeCodeCamp.com'
  }
};
```

以下是如何擷取物件屬性的值並將它們指定給同名變數的方法：

```js
const { johnDoe: { age, email }} = user;
```

以下是如何將物件屬性的值指定給不同名稱的變數：

```js
const { johnDoe: { age: userAge, email: userEmail }} = user;
```

# --instructions--

將兩個指派取代為等價的解構指派。它仍然應該將變數 `lowToday` 和 `highToday` 指定為來自 `LOCAL_FORECAST` 物件的 `today.low` 和 `today.high` 的值。

# --hints--

你應該移除 ES5 指派語法。

```js
assert(
  !__helpers.removeJSComments(code).match(/lowToday = LOCAL_FORECAST\.today\.low/g) &&
    !__helpers.removeJSComments(code).match(/highToday = LOCAL_FORECAST\.today.high/g)
);
```

你應該使用解構賦值來創建 `lowToday` 變數。

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*today\s*:\s*{\s*(low\s*:\s*lowToday[^}]*|[^,]*,\s*low\s*:\s*lowToday\s*)},?\s*}\s*=\s*LOCAL_FORECAST(;|\s+|\/\/)/g
  )
);
```

你應該使用解構賦值來創建 `highToday` 變數。

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*today\s*:\s*{\s*(high\s*:\s*highToday[^}]*|[^,]*,\s*high\s*:\s*highToday,?\s*)},?\s*}\s*=\s*LOCAL_FORECAST(;|\s+|\/\/)/g
  )
);
```

`lowToday` 應該相等於 `64`，而 `highToday` 應該相等於 `77`。

```js
assert(lowToday === 64 && highToday === 77);
```

# --seed--

## --seed-contents--

```js
const LOCAL_FORECAST = {
  yesterday: { low: 61, high: 75 },
  today: { low: 64, high: 77 },
  tomorrow: { low: 68, high: 80 }
};

// Only change code below this line
  
const lowToday = LOCAL_FORECAST.today.low;
const highToday = LOCAL_FORECAST.today.high;

// Only change code above this line
```

# --solutions--

```js
const LOCAL_FORECAST = {
  yesterday: { low: 61, high: 75 },
  today: { low: 64, high: 77 },
  tomorrow: { low: 68, high: 80 }
};
 
const { today: { low: lowToday, high: highToday }} = LOCAL_FORECAST;
```
