---
id: 587d7b87367417b2b2512b43
title: 使用箭頭函式撰寫簡潔的匿名函式（程式）
challengeType: 1
forumTopicId: 301211
dashedName: use-arrow-functions-to-write-concise-anonymous-functions
---

# --description--

在 JavaScript 中，我們經常不需要為函式命名，特別是在將函式作為引數傳遞給另一個函式時。相反地，我們會創建行內函式。因為我們不會在其他地方重複使用這些函式，所以不需要為它們命名。

為了達成此目的，我們經常使用以下語法：

```js
const myFunc = function() {
  const myVar = "value";
  return myVar;
}
```

ES6 提供了語法糖，讓我們不必這樣撰寫匿名函式（程式）。取而代之，你可以使用 **箭頭函式語法**：

```js
const myFunc = () => {
  const myVar = "value";
  return myVar;
}
```

當沒有函式（程式）主體，且只有返回值時，箭頭函式語法允許你省略關鍵字 `return` 以及包圍程式碼的括號。這有助於將較小的函式（程式）簡化為單行述語：

```js
const myFunc = () => "value";
```

這段程式碼預設仍會傳回字串 `value`。

# --instructions--

將指定給變數 `magic` 且傳回 `new Date()` 的函式改寫為箭頭函式語法。此外，確保沒有使用關鍵字 `var` 來定義任何東西。

# --hints--

你應該取代 `var` 關鍵字。

```js
assert.notMatch(code, /var/g)
```

`magic` 應該是常數變數（透過使用 `const`）。

```js
assert.match(code, /const\s+magic/g)
```

`magic` 應該是一個 `function`。

```js
assert(typeof magic === 'function');
```

`magic()` 應該傳回正確的日期。

```js
assert(magic().setHours(0, 0, 0, 0) === new Date().setHours(0, 0, 0, 0));
```

不應該使用 `function` 關鍵字。

```js
assert.notMatch(code, /function/g)
```

# --seed--

## --seed-contents--

```js
var magic = function() {
  return new Date();
};
```

# --solutions--

```js
const magic = () => {
  return new Date();
};
```
