---
id: 587d7b8a367417b2b2512b4d
title: 使用解構指派將物件作為函式（程式）的參數傳遞
challengeType: 1
forumTopicId: 301217
dashedName: use-destructuring-assignment-to-pass-an-object-as-a-functions-parameters
---

# --description--

在某些情況下，你可以直接在函式引數中解構物件。

請參考以下程式碼：

```js
const profileUpdate = (profileData) => {
  const { name, age, nationality, location } = profileData;

}
```

這會有效地解構傳入函式的物件。這也可以在原地完成：

```js
const profileUpdate = ({ name, age, nationality, location }) => {

}
```

當將 `profileData` 傳遞給上述函式時，值會從函式參數中解構以供函式內使用。

# --instructions--

在傳入 `half` 函式（程式）的引數中使用解構指派，只傳送 `max` 和 `min` 到函式內。

# --hints--

`stats` 應該是一個 `object`。

```js
assert(typeof stats === 'object');
```

`half(stats)` 應該是 `28.015`

```js
assert(half(stats) === 28.015);
```

應該使用解構賦值。

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/half=\({\w+,\w+}\)/));
```

應該使用解構的參數。

```js
assert(!__helpers.removeJSComments(code).match(/stats\.max|stats\.min/));
```

# --seed--

## --seed-contents--

```js
const stats = {
  max: 56.78,
  standard_deviation: 4.34,
  median: 34.54,
  mode: 23.87,
  min: -0.75,
  average: 35.85
};

// Only change code below this line
const half = (stats) => (stats.max + stats.min) / 2.0; 
// Only change code above this line
```

# --solutions--

```js
const stats = {
  max: 56.78,
  standard_deviation: 4.34,
  median: 34.54,
  mode: 23.87,
  min: -0.75,
  average: 35.85
};

const half = ( {max, min} ) => (max + min) / 2.0;
```
