---
id: 598f48a36c8c40764b4e52b3
title: 防止物件變異
challengeType: 1
forumTopicId: 301207
dashedName: prevent-object-mutation
---

# --description--

如同在前一個挑戰中所見，單獨使用 `const` 宣告並不能真正保護你的資料不被變更。為了確保你的資料不會改變，JavaScript 提供了一個函式 `Object.freeze` 來防止資料變更。

任何嘗試更改該物件都會被拒絕，且如果腳本在嚴格模式下執行，將會發出錯誤。

```js
let obj = {
  name:"FreeCodeCamp",
  review:"Awesome"
};
Object.freeze(obj);
obj.review = "bad";
obj.newProp = "Test";
console.log(obj); 
```

`obj.review` 和 `obj.newProp` 的指派會導致錯誤，因為我們的編輯器預設以嚴格模式執行，主控台將顯示值 `{ name: "FreeCodeCamp", review: "Awesome" }`。

# --instructions--

在這個挑戰中，你將使用 `Object.freeze` 來防止數學常數被更改。你需要凍結 `MATH_CONSTANTS` 物件，使沒有人能夠更改 `PI` 的值、添加或刪除屬性。

# --hints--

你不應該取代 `const` 關鍵字。

```js
assert(__helpers.removeJSComments(code).match(/const/g));
```

`MATH_CONSTANTS` 應該是一個常數變數（透過使用 `const`）。

```js
assert(__helpers.removeJSComments(code).match(/const\s+MATH_CONSTANTS/g));
```

你不應該更改 `MATH_CONSTANTS` 的原始宣告。

```js
assert(__helpers.removeJSComments(code).match(
   /const\s+MATH_CONSTANTS\s+=\s+{\s+PI:\s+3.14\s+};/g
));
```

`PI` 應該相等於 `3.14`。

```js
assert(PI === 3.14);
```

# --seed--

## --seed-contents--

```js
function freezeObj() {
  const MATH_CONSTANTS = {
    PI: 3.14
  };
  // Only change code below this line


  // Only change code above this line
  try {
    MATH_CONSTANTS.PI = 99;
  } catch(ex) {
    console.log(ex);
  }
  return MATH_CONSTANTS.PI;
}
const PI = freezeObj();
```

# --solutions--

```js
function freezeObj() {
  const MATH_CONSTANTS = {
    PI: 3.14
  };
  Object.freeze(MATH_CONSTANTS);

  try {
    MATH_CONSTANTS.PI = 99;
  } catch(ex) {
    console.log(ex);
  }
  return MATH_CONSTANTS.PI;
}
const PI = freezeObj();
```
