---
id: 5cdafbe72913098997531682
title: 使用 catch 處理被拒絕的 Promise
challengeType: 1
forumTopicId: 301204
dashedName: handle-a-rejected-promise-with-catch
---

# --description--

`catch` 是當你的 promise 被拒絕時使用的方法。它會在 promise 的 `reject` 方法被呼叫後立即執行。語法如下：

```js
myPromise.catch(error => {
  
});
```

`error` 是傳遞給 `reject` 方法的引數。

# --instructions--

為你的 promise 添加 `catch` 方法。使用 `error` 作為其回呼函式的參數，並將 `error` 記錄到主控台。

# --hints--

你應該在 promise 上呼叫 `catch` 方法。

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/(makeServerRequest|\))\.catch\(/g)
);
```

你的 `catch` 方法應該有一個以 `error` 為參數的回呼函式。

```js
assert(errorIsParameter);
```

你應該將 `error` 記錄到主控台。

```js
assert(
  errorIsParameter &&
    __helpers
      .removeWhiteSpace(__helpers.removeJSComments(code))
      .match(/\.catch\(.*?error.*?console.log\(error\).*?\)/)
);
```

# --seed--

## --after-user-code--

```js
const errorIsParameter = /\.catch\((function\(error\){|error|\(error\)=>)/.test(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)));
```

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to false to represent an unsuccessful response from a server
  let responseFromServer = false;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to false to represent an unsuccessful response from a server
  let responseFromServer = false;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});

makeServerRequest.catch(error => {
  console.log(error);
});
```
