---
id: 5cdafbd72913098997531681
title: 使用 then 處理已完成的 Promise
challengeType: 1
forumTopicId: 301203
dashedName: handle-a-fulfilled-promise-with-then
---

# --description--

當你的程式碼中有一個需要不確定時間的處理程序（例如非同步的東西），通常是伺服器請求時，Promises 最有用。當你發出伺服器請求時，會花費一些時間，且完成後你通常想對伺服器的回應做些事情。這可以透過使用 `then` 方法來達成。

```js
Promise.prototype.then(onFulfilled, onRejected)
```

`then` 方法會為 Promise 最終完成（無論是成功或拒絕）排程回呼函式。`onFulfilled` 和 `onRejected` 處理常式其中之一會被執行，以控制目前 Promise 的成功或拒絕。當 Promise 以 `resolve` 成功時，會呼叫 `onFulfilled` 處理常式。

```js
myPromise.then(result => {
  
});
```

`result` 來自傳給 `resolve` 方法的引數。

# --instructions--

為你的 promise 添加 `then` 方法，使用 `result` 作為其回呼函式的參數，並將 `result` 記錄到主控台。

# --hints--

你應該在 `promise` 上呼叫 `then` 方法。

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/(makeServerRequest|\))\.then\(/g)
);
```

你的 `then` 方法應該有一個以 `result` 為參數的回呼函式（程式）。

```js
assert(resultIsParameter);
```

你應該將 `result` 記錄到主控台。

```js
assert(
  resultIsParameter &&
    __helpers
      .removeWhiteSpace(__helpers.removeJSComments(code))
      .match(/\.then\(.*?result.*?console.log\(result\).*?\)/)
);
```

# --seed--

## --after-user-code--

```js
const resultIsParameter = /\.then\((function\(result\){|result|\(result\)=>)/.test(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)));
```

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to true to represent a successful response from a server
  let responseFromServer = true;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to true to represent a successful response from a server
  let responseFromServer = true;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});
```
