---
id: 5cdafbc32913098997531680
title: 使用 resolve 和 reject 完成一個 Promise
challengeType: 1
forumTopicId: 301196
dashedName: complete-a-promise-with-resolve-and-reject
---

# --description--

一個 promise 有三種狀態：`pending`、`fulfilled` 和 `rejected`。你在上一個挑戰中創建的 promise 永遠停留在 `pending` 狀態，因為你沒有添加完成 promise 的方式。給 promise 引數的 `resolve` 和 `reject` 參數用來完成這個動作。當你想讓 promise 成功時使用 `resolve`，當你想讓它失敗時使用 `reject`。這些是會帶入引數的方法，如下所示。

```js
const myPromise = new Promise((resolve, reject) => {
  if(condition here) {
    resolve("Promise was fulfilled");
  } else {
    reject("Promise was rejected");
  }
});
```

上面的範例使用字串作為這些函式的引數，但它實際上可以是任何東西。通常，它可能是一個物件，你會使用該物件中的資料，將其放在你的網站或其他地方。

# --instructions--

讓 promise 控制代碼成功和失敗。如果 `responseFromServer` 是 `true`，呼叫 `resolve` 方法以成功完成 promise。傳遞給 `resolve` 一個值為 `We got the data` 的字串。如果 `responseFromServer` 是 `false`，則改用 `reject` 方法並傳遞字串：`Data not received`。

# --hints--

當 `if` 條件為 `true` 時，應該以預期的字串呼叫 `resolve`。

```js
assert(
  __helpers.removeJSComments(code).match(/if\s*\(\s*responseFromServer\s*\)\s*{\s*resolve\s*\(\s*('|"|`)We got the data\1\s*\)(\s*|\s*;\s*)}/g)
);
```

當 `if` 條件為 `false` 時，應該以預期的字串呼叫 `reject`。

```js
assert(
  __helpers.removeJSComments(code).match(/}\s*else\s*{\s*reject\s*\(\s*('|"|`)Data not received\1\s*\)(\s*|\s*;\s*)}/g)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer represents a response from a server
  let responseFromServer;
    
  if(responseFromServer) {
    // Change this line
  } else {  
    // Change this line
  }
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer represents a response from a server
  let responseFromServer;

  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});
```
