---
id: 5cdafbb0291309899753167f
title: 創建一個 JavaScript Promise
challengeType: 1
forumTopicId: 301197
dashedName: create-a-javascript-promise
---

# --description--

JavaScript 中的 promise 就如其字面意義－你使用它來承諾執行某件事，通常是非同步的。當任務完成時，你要麼履行你的承諾，要麼未能履行。`Promise` 是一個建構式函式，因此你需要使用 `new` 關鍵字來創建一個。它接受一個函式作為引數，該函式有兩個參數－`resolve` 和 `reject`。這些是用來決議 promise 結果的方法。語法如下：

```js
const myPromise = new Promise((resolve, reject) => {

});
```

# --instructions--

建立一個名為 `makeServerRequest` 的新 promise。將一個帶有 `resolve` 和 `reject` 參數的函式傳入建構式。

# --hints--

你應該將一個 promise 指定給名為 `makeServerRequest` 的已宣告變數。

```js
assert(makeServerRequest instanceof Promise);
```

你的 promise 應該接收一個帶有 `resolve` 和 `reject` 參數的函式（程式）。

```js
assert(
  __helpers.removeJSComments(code).match(
    /Promise\s*\(\s*(function\s*\(\s*resolve\s*,\s*reject\s*\)\s*{|\(\s*resolve\s*,\s*reject\s*\)\s*=>\s*{)[^}]*}/g
  )
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {

});
```
