---
id: 5cddbfd622f1a59093ec611d
title: 創建模組腳本
challengeType: 6
forumTopicId: 301198
dashedName: create-a-module-script
---

# --description--

JavaScript 最初在主要是 HTML 的網頁中扮演著小角色。如今，它變得非常重要，有些網站幾乎完全使用 JavaScript 建立。為了讓 JavaScript 更加模組化、乾淨且易於維護，ES6 引入了一種方法，可以輕鬆地在 JavaScript 檔案之間共享程式碼。這涉及將檔案的部分內容匯出，以供一個或多個其他檔案使用，並在需要的地方匯入所需的部分。為了利用這項功能，你需要在你的 HTML 文件中創建一個 `type` 為 `module` 的腳本。以下是一個範例：

```html
<script type="module" src="filename.js"></script>
```

使用此 `module` 型別的腳本現在可以使用你將在接下來的挑戰中學習的 `import` 和 `export` 功能。

# --instructions--

將一個型別為 `module` 的腳本添加到 HTML 文件中，並將其原始碼檔案設為 `index.js`

# --hints--

你應該創建一個 `script` 頁籤。

```js
assert(__helpers.removeJSComments(code).match(/<\s*script[^>]*>\s*<\/\s*script\s*>/g));
```

你的 `script` 頁籤應該有 `type` 屬性且其值為 `module`。

```js
assert(
  __helpers.removeJSComments(code).match(
    /<\s*script\s+[^t]*type\s*=\s*('|")module\1[^>]*>\s*<\/\s*script\s*>/g
  )
);
```

你的 `script` 頁籤應該有一個 `src` 為 `index.js`。

```js
assert(
  __helpers.removeJSComments(code).match(
    /<\s*script\s+[^s]*src\s*=\s*('|")index\.js\1[^>]*>\s*<\/\s*script\s*>/g
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <!-- Only change code below this line -->

    <!-- Only change code above this line -->
  </body>
</html>
```

# --solutions--

```html
<html>
  <body>
    <script type="module" src="index.js"></script>
  </body>
</html>
```
