---
id: 5cfa550e84205a357704ccb6
title: 使用解構指派從物件中擷取值
challengeType: 1
forumTopicId: 301216
dashedName: use-destructuring-assignment-to-extract-values-from-objects
---

# --description--

<dfn>解構指派</dfn> 是 ES6 中介紹的特殊語法，用於整齊地指定直接從物件取得的值。

請考慮以下 ES5 程式碼：

```js
const user = { name: 'John Doe', age: 34 };

const name = user.name;
const age = user.age;
```

`name` 會有字串 `John Doe` 的值，而 `age` 會有數字 `34`。

以下是使用 ES6 解構語法的等價指派述語：

```js
const { name, age } = user;
```

同樣地，`name` 的值會是字串 `John Doe`，而 `age` 的值會是數字 `34`。

在這裡，`name` 和 `age` 變數將會被創建並指定來自 `user` 物件的各自值。你可以看到這樣寫有多乾淨。

你可以從該物件中擷取任意多或少的值。

# --instructions--

將兩個指派取代為等價的解構指派。它仍然應該將變數 `today` 和 `tomorrow` 指派為來自 `HIGH_TEMPERATURES` 物件的 `today` 和 `tomorrow` 的值。

# --hints--

你應該移除 ES5 指派語法。

```js
assert(
  !__helpers.removeJSComments(code).match(/today\s*=\s*HIGH_TEMPERATURES\.(today|tomorrow)/g)
);
```

你應該使用解構賦值來創建 `today` 變數。

```js
assert(
  __helpers.removeJSComments(code).match(/(var|let|const)\s*{\s*(today[^}]*|[^,]*,\s*today)\s*}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g)
);
```

你應該使用解構賦值來創建 `tomorrow` 變數。

```js
assert(
  __helpers.removeJSComments(code).match(/(var|let|const)\s*{\s*(tomorrow[^}]*|[^,]*,\s*tomorrow)\s*}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g)
);
```

`today` 應該要與 `77` 相等，且 `tomorrow` 應該要與 `80` 相等。

```js
assert(today === 77 && tomorrow === 80);
```

# --seed--

## --seed-contents--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

// Only change code below this line

const today = HIGH_TEMPERATURES.today;
const tomorrow = HIGH_TEMPERATURES.tomorrow;

// Only change code above this line
```

# --solutions--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

const { today, tomorrow } = HIGH_TEMPERATURES;
```
