---
id: 66bba6fff611169359d9d36a
title: 创建一个保密邮件页面
challengeType: 25
dashedName: build-a-confidential-email-page
demoType: onClick
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有个 `id` 为 `email` 的 `main` 元素。
1. 你的 `#email` 元素应设置 `padding` 为 `50px`、上外边距为 `50px`、`width` 为 `500px`、`border` 为 `2px`。
1. 你的 `#email` 元素的总宽度，包括内边距和边框，应为 `500px`。
1. 在 `#email` 元素中，你应该有两个 `div` 元素，一个的 id 是 `confidential`，另一个的 id 是 `top-secret`。
1. 你的 `#confidential` 和 `#top-secret` 元素都应该有个 `display`，值为 `inline-block`。
1. 你的 `#confidential` 和 `#top-secret` 元素都应该有 `padding`、左外边距和 `border`。
1. `#confidential` 元素的文本应为 `CONFIDENTIAL`。
1. `#top-secret` 元素的文本应为 `TOP SECRET`。
1. 你的 `#confidential` 和 `#top-secret` 元素应使用 CSS 变化进行旋转。
1. 你的 `#email` 元素中至少应有三个段落元素。
1. 在你的段落元素中，应至少有三个类为 `blurred` 的 `span` 元素。
1. 你应该有个 `blurred` 选择器，使用 CSS 滤镜将元素模糊 `3px`。

# --hints--

你应该有个 `id` 为 `email` 的 `main` 元素。

```js
assert.exists(document.querySelector('main#email'));
```

你应该有个 `#email` 选择器，将其元素所有边的 `padding` 设置为 `50px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('padding'), '50px');
```

你应该有个 `#email` 选择器，将其元素的 `margin-top` 设置为 `50px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('margin-top'), '50px');
```

你应该有个 `#email` 选择器，将其元素的 `width` 设置为 `500px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('width'), '500px');
```

你的 `#email` 元素应有 `2px` 的边框。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('border-width'), '2px');
```

你的 `#email` 元素应该有个 `box-sizing`，值为 `border-box`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.boxSizing, 'border-box');
```

在 `#email` 元素中，你应该有个 `id` 为 `confidential` 的 `div` 元素。

```js
assert.exists(document.querySelector('#email > div#confidential'));
```

在 `#email` 元素中，你应该有个 `id` 为 `top-secret` 的 `div` 元素。

```js
assert.exists(document.querySelector('#email > div#top-secret'));
```

你的 `#confidential` 元素应将其 `display` 设置为 `inline-block`。

```js
assert.equal(window.getComputedStyle(document.querySelector('#confidential'))?.display, 'inline-block');
```

你的 `#top-secret` 元素应将其 `display` 设置为 `inline-block`。

```js
assert.equal(window.getComputedStyle(document.querySelector('#top-secret'))?.display, 'inline-block');
```

你的 `#confidential` 元素应在四边都有 `padding`。

```js
const style = window.getComputedStyle(document.querySelector('#confidential'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

你的 `#top-secret` 元素应在四边都有 `padding`。

```js
const style = window.getComputedStyle(document.querySelector('#top-secret'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

你的 `#confidential` 元素应有一个左外边距。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.marginLeft, '0px');
```

你的 `#top-secret` 元素应有一个左外边距。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.marginLeft, '0px');
```

你的 `#confidential` 元素应有一个 `border`。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.borderWidth, '0px');
```

你的 `#top-secret` 元素应有一个 `border`。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.borderWidth, '0px');
```

你的 `#confidential` 元素的文本应为 `CONFIDENTIAL`。 你要么文本有遗漏，要么拼写有误。

```js
assert.equal(document.querySelector('#confidential').innerText.toLowerCase(), 'confidential' );
```

你的 `#top-secret` 元素的文本应为 `TOP SECRET`。 你要么文本有遗漏，要么拼写有误。

```js
assert.equal(document.querySelector('#top-secret').innerText.toLowerCase(), 'top secret' );
```

你的 `#confidential` 元素应使用 CSS `transform` 旋转该元素。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.transform, 'none');
```

你的 `#top-secret` 元素应使用 CSS `transform` 旋转该元素。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.transform, 'none');
```

你的 `#email` 元素中至少应有三个段落元素。

```js
assert.isAtLeast(document.querySelectorAll('#email > p').length, 3);
```

在你的段落元素中，应至少有三个类为 `blurred` 的 `span` 元素。

```js
assert.isAtLeast(document.querySelectorAll('p > span.blurred').length, 3);
```

你的 `.blurred` 元素不应为空。

```js
const els = document.querySelectorAll('.blurred');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.isAtLeast(el.innerText.length , 1))
```

你应有个 `.blurred` 选择器，将元素的 `filter` 设置为 `blur(3px)`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.blurred')?.filter, 'blur(3px)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Confidential Email</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Secret Marshmallow Mission</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="email">
    <div id="confidential">CONFIDENTIAL</div>
    <p>Dear Agent <span class="blurred">S'more,</span></p>
    <p>We have an emergency. The secret formula for our <span class="blurred">Mega Marshmallow</span> has been
      compromised. This formula is what makes our marshmallows the fluffiest and most delicious.</p>
    <p>We suspect that <span class="blurred">Professor Puff</span> is behind this. He has taken the formula to his
      hidden laboratory. Your mission is to infiltrate the lab and secure the formula before it's too late.</p>
    <p>Be sure to keep the lab's location confidential. Any leak of this information could jeopardize the entire
      operation.</p>
    <div id="top-secret">TOP SECRET</div>
  </main>
</body>

</html>
```

```css
#email {
  background-color: SeaShell;
  margin: 50px auto;
  padding: 50px;
  width: 500px;
  border: 2px solid black;
  box-shadow: 5px 3px 3px gray;
  box-sizing: border-box;
}

#confidential {
  display: inline-block;
  margin-left: 100px;
  padding: 10px;
  text-align: center;
  font-size: 30px;
  transform: rotate(-20deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

#top-secret {
  margin-left: 200px;
  display: inline-block;
  padding: 10px;
  text-align: center;
  font-size: 20px;
  transform: rotate(15deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

.blurred {
  filter: blur(3px);
}
```
