---
id: 587d8259367417b2b2512c85
title: "implementare l'algoritmo di ordinamento selection sort"
challengeType: 26
dashedName: implement-the-selection-sort-algorithm
---

# --description--

In questo laboratorio implementerai l'algoritmo di ordinamento selection sort. Il selection sort funziona selezionando il valore minimo in una lista e scambiandolo con il primo valore della lista. Poi inizia dalla seconda posizione, seleziona il valore più piccolo nella lista rimanente e lo scambia con il secondo elemento. Continua a iterare attraverso la lista e a scambiare gli elementi finché non raggiunge la fine della lista. Ora la lista è ordinata. Il selection sort ha complessità temporale quadratica in tutti i casi.

**Obiettivo:** soddisfa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione `selectionSort`.
1. La funzione `selectionSort` dovrebbe prendere un array e restituire un array con gli stessi elementi ma ordinati dal più piccolo al più grande.

# --hints--

`selectionSort` dovrebbe essere una funzione.

```js
assert.isFunction(selectionSort);
```

`selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` dovrebbe restituire un array che non cambia tranne per l'ordine.

```js
assert.sameMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`selectionSort` dovrebbe restituire un array ordinato (dal più piccolo al più grande).

```js
assert.sameOrderedMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
)
```

`selectionSort` non dovrebbe usare il metodo integrato `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
