---
id: 6140c7e645d8e905819f1dd4
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

لبداية، أضف عنصر `link` لملف `./styles.css`.

# --hints--

يجب أن تحتوي شفرتك على عنصر `link`.

```js
assert.match(code, /<link/);
```

يجب أن يكون لديك عنصر `link` واحد.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

يجب أن يكون عنصر `link` داخل عنصر `head`.

```js
assert.exists(document.querySelector('head > link'));
```

يجب أن يحتوي عنصر `link` على خاصية `rel` بالقيمة `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

يجب أن يحتوي عنصر `link` على خاصية `href` بالقيمة `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>

  </body>
</html>

```

```css

```
