---
id: 6140c9d35015ae0ba0c250e8
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

أضف `div` داخل عنصر `body` وأعطه `class` باسم `wheel`.

داخل `div` الجديد، أضف ستة عناصر `span` مع تعيين `class` إلى `line`، وستة عناصر `div` مع تعيين `class` إلى `cabin`.

# --hints--

يجب أن تنشئ `div` جديد داخل عنصر `body`.

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert.lengthOf(divs, 1);
```

يجب أن يحتوي `div` داخل عنصر `body` على `class` باسم `wheel`.

```js
assert.isTrue(document.querySelector('body div')?.classList?.contains('wheel'));
```

يجب أن يحتوي عنصر `.wheel` على ستة عناصر `span` بداخله.

```js
assert.lengthOf(document.querySelectorAll('.wheel span'), 6);
```

يجب أن تحتوي ستة عناصر `span` داخل عنصر `.wheel` على `class` باسم `line`.

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert.isTrue(spans?.every(span => span?.classList?.contains('line')));
assert.lengthOf(document.querySelectorAll('.line'), 6);
```

يجب أن يحتوي عنصر `.wheel` على ستة عناصر `div` بداخله.

```js
assert.lengthOf(document.querySelectorAll('.wheel div'), 6);
```

يجب أن تحتوي ستة عناصر `div` داخل عنصر `.wheel` على `class` باسم `cabin`.

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert.isTrue(divs?.every(div => div?.classList?.contains('cabin')));
assert.lengthOf(document.querySelectorAll('.cabin'), 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css

```
