---
id: 6140cd32d018ed0f600eefce
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

عيّن خاصية `position` لمحدِّد `.wheel` إلى `absolute`. عيّن خاصيتي `height` و `width` كلاهما إلى `55vw`.

# --hints--

يجب أن يحتوي محدِّد `.wheel` على خاصية `position` مضبوطة على `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.position, 'absolute');
```

يجب أن يحتوي محدِّد `.wheel` على خاصية `height` مضبوطة على `55vw`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.height, '55vw');
```

يجب أن يحتوي محدِّد `.wheel` على خاصية `width` مضبوطة على `55vw`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.width, '55vw');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
