---
id: 6140d1a351e88f159ed54fca
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

باستخدام نفس النمط، أنشئ محدِّدًا منفصلًا للعنصر الثالث `.line`، والرابع `.line`، والخامس `.line`، والسادس `.line`.

عيّن خاصية `transform` للعنصر الثالث `.line` إلى `rotate(120deg)`، والرابع إلى `rotate(180deg)`، والخامس إلى `rotate(240deg)`، والسادس إلى `rotate(300deg)`.

# --hints--

يجب أن تنشئ محدِّد `.line:nth-of-type(3)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)'));
```

يجب أن يحتوي محدِّد `.line:nth-of-type(3)` على خاصية `transform` مضبوطة إلى `rotate(120deg)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)')?.transform, 'rotate(120deg)');
```

يجب أن تنشئ محدِّد `.line:nth-of-type(4)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)'));
```

يجب أن يحتوي محدِّد `.line:nth-of-type(4)` على خاصية `transform` مضبوطة إلى `rotate(180deg)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)')?.transform, 'rotate(180deg)');
```

يجب أن تنشئ محدِّد `.line:nth-of-type(5)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)'));
```

يجب أن يحتوي محدِّد `.line:nth-of-type(5)` على خاصية `transform` مضبوطة إلى `rotate(240deg)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)')?.transform, 'rotate(240deg)');
```

يجب أن تنشئ محدِّد `.line:nth-of-type(6)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)'));
```

يجب أن يحتوي محدِّد `.line:nth-of-type(6)` على خاصية `transform` مضبوطة إلى `rotate(300deg)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)')?.transform, 'rotate(300deg)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}

--fcc-editable-region--

--fcc-editable-region--
```
