---
id: 68ed661e9f463c155bb677a1
title: Hatua 1
challengeType: 0
dashedName: step-1
---

# --description--

Katika warsha hii utafanya mazoezi ya kutumia CSS flexbox kwa kubuni seti ya masanduku yenye rangi.

Kuanza, ndani ya kipengele cha `body`, ongeza kipengele cha `header`. Panga ndani kipengele cha `h1` ndani ya kipengele cha `header` chenye maandishi `Colored Boxes Layout`.

# --hints--

Unapaswa kuwa na kipengele cha `header` ndani ya kipengele cha `body`.

```js
const headerEl = document.querySelector('body > header');
assert.exists(headerEl);
```

Kipengele chako cha `header` kinapaswa kuwa na kipengele cha `h1` ndani yake.

```js
const h1El = document.querySelector('body > header > h1');
assert.exists(h1El);
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi `Colored Boxes Layout` ndani yake.

```js
const h1El = document.querySelector('body > header > h1');
assert.equal(h1El?.textContent.trim(), 'Colored Boxes Layout');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
</head>
<body>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```
