---
id: 68ef1b1c657f1b2396d98592
title: Hatua 29
challengeType: 0
dashedName: step-29
---

# --description--

Sifa ya `flex` inadhibiti ukubwa na tabia za vitu ndani ya kontena lenye kubadilika. Inajumuisha sifa tatu: `flex-grow`, `flex-shrink`, na `flex-basis`.

Sifa ya `flex-grow` inadhibiti ni nafasi gani ya ziada kipande cha kubadilika kinapaswa kuchukua ikiwa kuna nafasi ya bure ndani ya kontena.

Sasa, tukirudi kwenye kichaguzi cha `.box`, badilisha nambari ya kwanza katika `flex` (inayolingana na thamani ya `flex-grow`) kutoka `1` hadi `0`.

# --hints--

Kichaguzi chako cha `.box` kinapaswa kuwa na sifa ya `flex` yenye thamani ya `0 1 100px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('flex'), '0 1 100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Colored Boxes</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <header>
    <h1>Colored Boxes Layout</h1>
  </header>
  <div class="flex-container">
    <div class="box">
      <h2>Box 1</h2>
      <p>Red</p>
    </div>
    <div class="box">
      <h2>Box 2</h2>
      <p>Orange</p>
    </div>
    <div class="box">
      <h2>Box 3</h2>
      <p>Yellow</p>
    </div>
    <div class="box">
      <h2>Box 4</h2>
      <p>Green</p>
    </div>
    <div class="box">
      <h2>Box 5</h2>
      <p>Blue</p>
    </div>
    <div class="box">
      <h2>Box 6</h2>
      <p>Indigo</p>
    </div>
  </div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
--fcc-editable-region--
  flex: 1 1 100px;
--fcc-editable-region--
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0;
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

```
