---
id: 68ef0b3237fef991c92d9c19
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Sasa ongeza kipengele cha `div` chenye darasa la vitu la `flex-container` chini ya `header` yako. Kisha tengeneza vipengele sita vya `div` vyenye darasa la vitu la `box` ndani ya `div` uliyotengeneza sasa hivi.

Kumbuka kuwa sifa ya kipengele cha darasa la vitu hutumika mara nyingi kuashiria jina la darasa la vitu katika karatasi ya mitindo. Katika kesi hii, vipengele vyote vya `div` vitapambwa kwa usawa kulingana na ufafanuzi wa mitindo wa `.flex-container` na `.box`.

# --hints--

Unapaswa kuwa na kipengele cha `div` chini ya kipengele cha `header`.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.exists(divELBelowHeader);
```

Kipengele chako cha `div` chini ya kipengele cha `header` kinapaswa kuwa na darasa la vitu la `flex-container`.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.equal(divELBelowHeader.className, 'flex-container');
```

Unapaswa kuwa na vipengele sita vya `div` vilivyopangwa ndani ya kipengele cha `.flex-container`.

```js
const divElsInsideContainer = document.querySelectorAll('div.flex-container > div');
assert.lengthOf(divElsInsideContainer, 6);
```

Vipengele vyote sita vya `div` ndani ya kipengele cha `.flex-container` vinapaswa kuwa na darasa la vitu la `box`.

```js
const divElsInsideContainerWithBox = document.querySelectorAll('div.flex-container > div.box');
assert.lengthOf(divElsInsideContainerWithBox, 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```

```css
h1 {
	text-align: center;
	margin-bottom: 10px;
}
```
