---
id: 68f91df5776e693049fb8942
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Sasa ni wakati wa kuongeza baadhi ya sheria za CSS katika faili la `styles.css`. Anza kwa kuunda kichaguzi kwa kipengele cha `h1`.

Weka maudhui ya kipengele cha `h1` katikati kwa kuweka sifa yake ya `text-align` kuwa `center`. Pia unda sifa ya `margin-bottom` yenye thamani ya `10px` ili kuongeza ukingo kati ya chini ya kipengele cha `h1` na kipengele chochote cha HTML kinachofuata chini.

# --hints--

Unapaswa kuwa na kichaguzi cha `h1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

Kichaguzi chako cha `h1` kinapaswa kuweka sifa ya `text-align` kuwa `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('text-align'), 'center');
```

Kichaguzi chako cha `h1` kinapaswa kuweka sifa ya `margin-bottom` kuwa `10px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('margin-bottom'), '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
