---
id: 692f314d449329ba8bcd5d2e
title: Hatua 17
challengeType: 0
dashedName: step-17
---

# --description--

Sifa ya `border-radius` inazungusha pembe za mpaka za kipengee. Unaweza kubainisha:

- Thamani moja itakayotumika kwa pembe zote nne.
- Thamani mbili, ambapo thamani ya kwanza ni kwa pembe ya juu-kushoto/chini-kulia, na ya pili kwa pembe ya juu-kulia/chini-kushoto.
- Thamani tatu, zinazolingana na pembe ya juu-kushoto, juu-kulia/chini-kushoto, chini-kulia.
- Thamani nne, ambazo huweka mduara wa mpaka kwa mwelekeo wa saa kuanzia pembe ya juu-kushoto.

Sasa ongeza sifa `border-radius` yenye thamani `5px` kwa kichaguzi `.box`.

# --hints--

Kichaguzi chako cha `.box` kinapaswa kuwa na sifa `border-radius` yenye thamani `5px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('border-radius'), '5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
