---
id: 692f22933ba1cb86cd75a70c
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Sasa ni wakati wa kuunda jina la rangi kwa masanduku. Ili kufanya hivyo, ongeza kipengele cha `p` chini ya kila kipengele cha `h2` kwa majina haya ya rangi kwa mpangilio: `Red`, `Orange`, `Yellow`, `Green`, `Blue`, na `Indigo`. Rangi moja kwa kila sanduku.

# --hints--

Unapaswa kuongeza kipengele cha `p` chenye maandishi ya `Red` chini ya kipengele cha `h2` ndani ya kipengele cha kwanza cha `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[0];
assert.equal(heading.textContent, "Red");
```

Unapaswa kuongeza kipengele cha `p` chenye maandishi ya `Orange` chini ya kipengele cha `h2` ndani ya kipengele cha pili cha `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[1];
assert.equal(heading.textContent, "Orange");
```

Unapaswa kuongeza kipengele cha `p` chenye maandishi ya `Yellow` chini ya kipengele cha `h2` ndani ya kipengele cha tatu cha `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[2];
assert.equal(heading.textContent, "Yellow");
```

Unapaswa kuongeza kipengele cha `p` chenye maandishi ya `Green` chini ya kipengele cha `h2` ndani ya kipengele cha nne cha `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[3];
assert.equal(heading.textContent, "Green");
```

Unapaswa kuongeza kipengele cha `p` chenye maandishi ya `Blue` chini ya kipengele cha `h2` ndani ya kipengele cha tano cha `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[4];
assert.equal(heading.textContent, "Blue");
```

Unapaswa kuongeza kipengele cha `p` chenye maandishi ya `Indigo` chini ya kipengele cha `h2` ndani ya kipengele cha sita cha `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[5];
assert.equal(heading.textContent, "Indigo");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
    <div class="box">
      <h2>Box 1</h2>
			
    </div>
    <div class="box">
      <h2>Box 2</h2>
			
    </div>
    <div class="box">
      <h2>Box 3</h2>
			
    </div>
    <div class="box">
      <h2>Box 4</h2>
			
    </div>
    <div class="box">
      <h2>Box 5</h2>
			
    </div>
    <div class="box">
      <h2>Box 6</h2>
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
