---
id: 692f2c20cb7f9e5050b935a7
title: Hatua 9
challengeType: 0
dashedName: step-9
---

# --description--

Ongeza sifa `border` yenye thamani `1px solid #000` kwa kichaguzi `.box`. Sifa ya `border` hutumika kuunda mpangilio unaoonekana kuzunguka kipengele. Inaweza kubinafsishwa kwa upande wa `border-width`, `border-style`, na `border-color`.

Katika kesi hii, `1px` ni `border-width`, ambayo inaelezea unene wa mpaka. `solid` ni `border-style` na huweka muonekano wa mpaka (imara, yenye mistari, yenye nukta na kadhalika). Na `#000` ni `border-color`, ambayo inaelezea rangi ya mpaka.

# --hints--

Kichaguzi chako cha `.box` kinapaswa kuwa na sifa `border` yenye thamani `1px solid #000`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('border'), '1px solid rgb(0, 0, 0)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Colored Boxes</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <header>
    <h1>Colored Boxes Layout</h1>
  </header>
  <div class="flex-container">
    <div class="box">
      <h2>Box 1</h2>
      <p>Red</p>
    </div>
    <div class="box">
      <h2>Box 2</h2>
      <p>Orange</p>
    </div>
    <div class="box">
      <h2>Box 3</h2>
      <p>Yellow</p>
    </div>
    <div class="box">
      <h2>Box 4</h2>
      <p>Green</p>
    </div>
    <div class="box">
      <h2>Box 5</h2>
      <p>Blue</p>
    </div>
    <div class="box">
      <h2>Box 6</h2>
      <p>Indigo</p>
    </div>
  </div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
