---
id: 692f22a2ad44ae887a24c8f9
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Ndani ya kila kipengele cha `.box`, ongeza kipengele cha `h2` chenye maandishi ya `Box` ikifuatiwa na nafasi na nambari mfululizo kuanzia `1`.

# --hints--

Unapaswa kuwa na kipengele cha `h2` chenye maandishi ya `Box 1` ndani ya `div` ya kwanza yenye darasa la `box`.

```js
const firstH2El = document.querySelector("div.box:nth-of-type(1) > h2");
assert.equal(firstH2El.textContent, "Box 1");
```

Unapaswa kuwa na kipengele cha `h2` chenye maandishi ya `Box 2` ndani ya `div` ya pili yenye darasa la `box`.

```js
const secondH2El = document.querySelector("div.box:nth-of-type(2) > h2");
assert.equal(secondH2El.textContent, "Box 2");
```

Unapaswa kuwa na kipengele cha `h2` chenye maandishi ya `Box 3` ndani ya `div` ya tatu yenye darasa la `box`.

```js
const thirdH2El = document.querySelector("div.box:nth-of-type(3) > h2");
assert.equal(thirdH2El.textContent, "Box 3");
```

Unapaswa kuwa na kipengele cha `h2` chenye maandishi ya `Box 4` ndani ya `div` ya nne yenye darasa la `box`.

```js
const fourthH2El = document.querySelector("div.box:nth-of-type(4) > h2");
assert.equal(fourthH2El.textContent, "Box 4");
```

Unapaswa kuwa na kipengele cha `h2` chenye maandishi ya `Box 5` ndani ya `div` ya tano yenye darasa la `box`.

```js
const fifthH2El = document.querySelector("div.box:nth-of-type(5) > h2");
assert.equal(fifthH2El.textContent, "Box 5");
```

Unapaswa kuwa na kipengele cha `h2` chenye maandishi ya `Box 6` ndani ya `div` ya mwisho yenye darasa la `box`.

```js
const sixthH2El = document.querySelector("div.box:nth-of-type(6) > h2");
assert.equal(sixthH2El.textContent, "Box 6");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Colored Boxes</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <header>
    <h1>Colored Boxes Layout</h1>
  </header>
  <div class="flex-container">
    <div class="box">
--fcc-editable-region--
      
    </div>
    <div class="box">
      
    </div>
    <div class="box">
      
    </div>
    <div class="box">
      
    </div>
    <div class="box">
      
    </div>
    <div class="box">
      
--fcc-editable-region--
    </div>
  </div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

```
