---
id: 69899123889dcf5b19be1847
title: 实现深度优先搜索算法
challengeType: 26
dashedName: implement-depth-first-search-algorithm
---

# --description--

在本实验中，你将实现一种称为 <dfn>深度优先搜索</dfn>（DFS）的图遍历算法。

当 <dfn>广度优先</dfn>（BFS）按层级逐层探索邻居时，<dfn>深度优先搜索</dfn> 会尽可能深入单一路径的边，然后再回溯。

一旦算法到达路径的末端（一个没有未访问邻居的节点），它会回溯到最近的仍有未探索边的节点，并从那里继续查找。

要迭代地实现这个算法，你需要使用一个**栈**。在 JavaScript 中，标准的 `Array` 通过 `.push()` 和 `.pop()` 方法完美地作为栈使用，遵循**后进先出（LIFO）**原则。这确保了最近添加到栈中的邻居是下一个被探索的。

该算法的简单输出是可从给定节点到达的节点列表。 因此，在实施此算法时，你需要跟踪你已经访问过的节点。

**目标**：完成以下用户故事并通过所有测试以完成实验。

**用户故事：**

1. 你应该定义一个名为 `dfs` 的函数。
1. `dfs` 函数应接受两个参数：

    - `graph`：一个邻接矩阵（一个数组的数组）。
    - `root`：一个数字节点标签（起始索引），其数值为介于 `0` 和 `n - 1` 之间的节点数字，其中 `n` 是图中节点的总数。

1. 该函数应返回一个包含从起始节点可达的所有节点标签的数组。

# --hints--

你应该有一个名为 `dfs` 的函数，它接受两个参数。

```js
assert.isFunction(dfs);
assert.lengthOf(dfs, 2);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)` 应返回包含 `1`、`2`、`3` 和 `0` 的列表。

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)` 应该返回一个包含 `1`、`2`、`3` 和 `0` 的列表。

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)` 应该返回 `[3]`。

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3), [3]);
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)` 应返回一个包含 `3` 和 `2` 的列表。

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3), [3, 2]);

```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)` 应返回包含 `0` 和 `1` 的列表。

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0), [0, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfs(graph, root) {
  let stack = [root];
  let visited = [];

  while (stack.length > 0) {
    let current = stack.pop();

    if (!visited.includes(current)) {
      visited.push(current);
      let neighbors = graph[current];
      
      for (let i = 0; i < neighbors.length; i++) {
        if (neighbors[i] === 1) {
          stack.push(i);
        }
      }
    }
  }
  return visited;
}
```
