---
id: 5ddb965c65d27e1512d44d9a
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このプロジェクトでは、カロリー計算機フォームを作成します。このフォームには、ユーザーが一日の摂取予定カロリーとさまざまな食事の摂取カロリーを入力できます。 そして、ユーザーがカロリー不足か過剰かを計算して表示します。

ボイラープレートの CSS と HTML はあらかじめ提供されています。 しかし、カロリー計算機のフォーム部分はあなたの手で構築する必要があります。

HTML と CSS のコードを自由に閲覧してください。その後、`form` 要素を 1 つ追加して、`id` を `calorie-counter` に設定してください。

# --hints--

`form` 要素が必要です。

```js
assert.exists(document.querySelector('form'));
```

あなたの `form` 要素には `id` 属性があり、`calorie-counter` に設定されている必要があります。

```js
assert.equal(document.querySelector('form')?.id, 'calorie-counter');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
