---
id: 63b6075a62883218d282504c
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

`input` 要素にはいくつか追加の属性が必要です。 数値のみの入力を許可するため `type` 属性を `number` に、また、正の数のみを許可するため `min` 属性を `0` に設定してください。そして `placeholder` を `Daily calorie budget` に設定してください。

最後に、`input` 要素を `required` (必須) に設定してください。

# --hints--

`input` 要素の `type` 属性は `number` に設定してください。

```js
assert.equal(document.querySelector('form input').type, 'number');
```

`input` 要素の `min` 属性は `0` に設定してください。

```js
assert.equal(document.querySelector('form input').min, '0');
```

`input` 要素の `placeholder` 属性は `Daily calorie budget` に設定してください。

```js
assert.equal(document.querySelector('form input').placeholder, 'Daily calorie budget');
```

`input` 要素には `required` 属性が必要です。

```js
assert.equal(document.querySelector('form input').hasAttribute('required'), true);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
--fcc-editable-region--
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input id="budget" />
        </form>
--fcc-editable-region--
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
