---
id: 63b60af1a0b9f7238a9dd294
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

セレクトメニューに、先ほど作成した食事と運動の各 `fieldset` 要素に対応する選択肢を追加しましょう。 `option` 要素を使用して、各 `fieldset` に対応する選択肢を作成してください。 各オプションの `value` 属性は `fieldset` の `id` に、テキストは `legend` のテキストにしてください。

また、`Breakfast` オプションを `selected` (既定で選択状態) に設定してください。

# --hints--

`select` 要素内に `option` 要素を 5 つ作成してください。

```js
assert.equal(document.querySelectorAll('.controls select option')?.length, 5);
```

最初の `option` には `Breakfast` というテキストを入れてください。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[0]?.textContent?.trim(), 'Breakfast');
```

最初の `option` の `value` 属性を `breakfast` に設定してください。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[0]?.value, 'breakfast');
```

2 つ目の `option` には `Lunch` というテキストを入れてください。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[1]?.textContent?.trim(), 'Lunch');
```

2 つ目の `option` の `value` 属性を `lunch` に設定してください。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[1]?.value, 'lunch');
```

3 つ目の `option` には `Dinner` というテキストを入れてください。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[2]?.textContent?.trim(), 'Dinner');
```

3 つ目の `option` の `value` 属性を `dinner` に設定してください。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[2]?.value, 'dinner');
```

4 つ目の `option` には `Snacks` というテキストを入れてください。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[3]?.textContent?.trim(), 'Snacks');
```

4 つ目の `option` の `value` 属性を `snacks` に設定してください。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[3]?.value, 'snacks');
```

5 つ目の `option` には `Exercise` というテキストを入れてください。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[4]?.textContent?.trim(), 'Exercise');
```

5 つ目の `option` の `value` 属性を `exercise` に設定してください。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[4]?.value, 'exercise');
```

最初の `option` がデフォルトで選択されている必要があります。

```js
const isFirstOptionSelected = document.querySelectorAll('.controls select option')?.[0]?.getAttributeNames()?.includes('selected');
const selectedOptions = document.querySelectorAll('.controls select option[selected]');
assert.isTrue(isFirstOptionSelected && selectedOptions.length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--
          <div class="controls">
            <span>
              <label for="entry-dropdown">Add food or exercise:</label>
              <select id="entry-dropdown" name="options">

              </select>
              <button type="button" id="add-entry">Add Entry</button>
            </span>
          </div>
--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
