---
id: 63b606f09a14cc1781aea1fb
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

このフォームでは、ユーザーが自分の一日の摂取予定カロリーを表す数値を入力できるようにします。

まず `label` 要素を作成し、`for` 属性を `budget` に設定し、テキストは `Budget` としてください。次に `input` 要素を作成して、`id` を `budget` に設定してください。

# --hints--

`form` 内に `label` 要素を作成する必要があります。

```js
assert.exists(document.querySelector('form label'));
```

`label` 要素の `for` 属性は `budget` に設定してください。

```js
assert.equal(document.querySelector('form label')?.getAttribute('for'), 'budget');
```

`label` 要素がテキスト `Budget` を持つようにしてください。

```js
assert.equal(document.querySelector('form label')?.innerText, 'Budget');
```

`form` 内に `input` 要素を作成する必要があります。

```js
assert.exists(document.querySelector('form input'));
```

`input` 要素は `label` 要素の後に来る必要があります。

```js
assert.equal(document.querySelector('form input')?.previousElementSibling?.tagName, "LABEL");
```

`input` 要素の `id` は `budget` に設定してください。

```js
assert.equal(document.querySelector('form input')?.id, "budget");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
--fcc-editable-region--
        <form id="calorie-counter">

        </form>
--fcc-editable-region--
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
