---
id: 587d824f367417b2b2512c59
title: Запустіть функціональні тести кінцевих точок API за допомогою Chai-HTTP II
challengeType: 2
forumTopicId: 301592
dashedName: run-functional-tests-on-api-endpoints-using-chai-http-ii
---

# --description--

Нагадаємо, цей проєкт створюється на основі шаблонного проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

У межах `tests/2_functional-tests.js` змініть тест `'Test GET /hello with your name'` (`// #2`), щоб перевірити `status` і `text` відповіді, щоб тест пройшов успішно.

Надішліть своє ім’я як URL-запит, додавши `?name=<your_name>` до маршруту. Кінцева точка відповідає `'hello <your_name>'`.

# --hints--

Всі тести мають пройти

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=1');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Ви маєте перевірити, що `res.status` == 200

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=1');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

Ви маєте перевірити, що `res.text` == `'hello <your_name>'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=1');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.text');
assert.match(data.assertions[1].args[1], /hello [\w\d_-]/);
```
