---
id: 587d824c367417b2b2512c4f
title: "Перевірка, чи входить значення в певний діапазон"
challengeType: 2
forumTopicId: 301598
dashedName: test-if-a-value-falls-within-a-specific-range
---

# --description--

Нагадаємо, цей проєкт створюється на основі стартового проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

```javascript
.approximately(actual, expected, delta, [message])
```

Твердження, що `actual` дорівнює `expected` в межах діапазону +/- `delta`.

# --instructions--

У `tests/1_unit-tests.js` у тесті з міткою `#10` у наборі `Comparisons` змініть кожен `assert` на `assert.approximately`, щоб тест пройшов (повинен оцінюватися як `true`).

Виберіть мінімальний діапазон (3-й параметр), щоб тест завжди проходив. Він має бути меншим за 1.

# --hints--

Всі тести мають проходити.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Вам слід вибрати правильний діапазон для першого твердження — `approximately(actual, expected, range)`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'approximately');
assert.equal(
  data.assertions[0].args[2],
  0.5,
  "weirdNumbers(0.5) is in the range (0.5, 1.5]. It's within 1 +/- 0.5"
);
```

Вам слід вибрати правильний діапазон для другого твердження — `approximately(actual, expected, range)`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'approximately');
assert.equal(
  data.assertions[1].args[2],
  0.8,
  "weirdNumbers(0.2) is in the range (0.2, 1.2]. It's within 1 +/- 0.8"
);
```
