---
id: 587d824d367417b2b2512c54
title: Використання регулярних виразів для перевірки рядка
challengeType: 2
forumTopicId: 301608
dashedName: use-regular-expressions-to-test-a-string
---

# --description--

Нагадаємо, цей проєкт створюється на основі стартового проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`match()` перевіряє, чи відповідає фактичне значення регулярному виразу, переданому другим аргументом.

# --instructions--

У `tests/1_unit-tests.js` у тесті з міткою `#15` у наборі `Strings` замініть кожен `assert` на `assert.match` або `assert.notMatch`, щоб тест пройшов (повинен оцінюватися як `true`). Не змінюйте аргументи, передані твердженням.

# --hints--

Всі тести мають пройти.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=14');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Вам слід обрати правильний метод для першого твердження — `match` чи `notMatch`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=14');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'match',
  "'# name:John Doe, age:35' matches the regex"
);
```

Вам слід обрати правильний метод для другого твердження — `match` чи `notMatch`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=14');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'notMatch',
  "'# name:Paul Smith III, age:twenty-four' does not match the regex (the age must be numeric)"
);
```
