---
id: 587d824a367417b2b2512c46
title: "Дізнайтеся, як працюють твердження в JavaScript"
challengeType: 2
forumTopicId: 301589
dashedName: learn-how-javascript-assertions-work
---

# --description--

Робота над цими завданнями передбачає написання вашого коду з використанням одного з наступних методів:

- Клонуйте <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">цей репозиторій GitHub</a> і виконайте ці завдання локально.
- Використайте будь-який конструктор сайтів на свій вибір для виконання проєкту. Обов’язково включіть усі файли з нашого репозиторію GitHub.

# --instructions--

У `tests/1_unit-tests.js` у тесті з міткою `#1` у наборі `Basic Assertions` замініть кожен `assert` на `assert.isNull` або `assert.isNotNull`, щоб тест пройшов (повинен оцінюватися як `true`). Не змінюйте аргументи, передані твердженням.

# --hints--

Всі тести мають пройти.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Вам слід обрати правильний метод для першого твердження — `isNull` чи `isNotNull`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isNull', 'Null is null');
```

Вам слід обрати правильний метод для другого твердження — `isNull` чи `isNotNull`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'isNotNull', '1 is not null');
```
