---
id: 587d824b367417b2b2512c47
title: "Перевірка, чи визначена змінна або функція"
challengeType: 2
forumTopicId: 301602
dashedName: test-if-a-variable-or-function-is-defined
---

# --description--

Нагадаємо, цей проєкт створюється на основі стартового проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

У `tests/1_unit-tests.js` під тестом з міткою `#2` у наборі `Basic Assertions` замініть кожен `assert` на `assert.isDefined()` або `assert.isUndefined()`, щоб тест пройшов (повинен оцінюватися як `true`). Не змінюйте аргументи, передані твердженням.

# --hints--

Всі тести мають пройти.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Виберіть правильний метод для першого твердження — `isDefined` чи `isUndefined`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'isDefined',
    'Null is not undefined'
  );
```

Виберіть правильний метод для другого твердження — `isDefined` чи `isUndefined`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'isUndefined',
    'Undefined is undefined'
  );
```

Виберіть правильний метод для третього твердження — `isDefined` чи `isUndefined`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'isDefined',
    'A string is not undefined'
  );
```
