---
id: 587d824e367417b2b2512c58
title: Запустіть функціональні тести кінцевих точок API за допомогою Chai-HTTP
challengeType: 2
forumTopicId: 301593
dashedName: run-functional-tests-on-api-endpoints-using-chai-http
---

# --description--

Нагадаємо, цей проєкт створюється на основі стартового проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Mocha дозволяє тестувати асинхронні операції, такі як виклики кінцевих точок API, за допомогою плагіна `chai-http`.

Нижче наведено приклад тесту з використанням `chai-http` для набору тестів під назвою `'GET /hello?name=[name] => "hello [name]"'`:

```js
suite('GET /hello?name=[name] => "hello [name]"', function () {
  test('?name=John', function (done) {
    chai
      .request(server)
      .keepOpen()
      .get('/hello?name=John')
      .end(function (err, res) {
        assert.equal(res.status, 200, 'Response status should be 200');
        assert.equal(res.text, 'hello John', 'Response should be "hello John"');
        done();
      });
  });
});
```

Тест надсилає запит `GET` до сервера з ім’ям у рядку запиту URL (`?name=John`). У функції зворотного виклику методу `end` отримується об’єкт відповіді (`res`), який містить властивість `status`.

Перший `assert.equal` перевіряє, чи статус дорівнює `200`. Другий `assert.equal` перевіряє, що рядок відповіді (`res.text`) дорівнює `"hello John"`.

Також зверніть увагу на параметр `done` у функції зворотного виклику тесту. Викликати його без аргументів наприкінці тесту необхідно, щоб сигналізувати про завершення асинхронної операції.

Нарешті, зверніть увагу на метод `keepOpen` одразу після методу `request`. Зазвичай ви запускаєте свої тести з командного рядка або як частину автоматизованого процесу інтеграції, і ви можете дозволити `chai-http` автоматично запускати і зупиняти ваш сервер.

Однак тести, які запускаються при відправленні посилання на ваш проєкт, вимагають, щоб сервер був увімкнений, тому потрібно використовувати метод `keepOpen`, щоб запобігти зупинці сервера `chai-http`.

# --instructions--

У межах `tests/2_functional-tests.js` змініть тест `'Test GET /hello with no name'` (`// #1`), щоб перевірити `status` і `text` відповіді, щоб тест пройшов. Не змінюйте аргументи, передані в assert.

URL-запит бути відсутнім. Без імені у рядку запиту кінцева точка відповідає `hello Guest`.

# --hints--

Всі тести мають пройти

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Ви маєте перевірити, що `res.status` == 200

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

Ви маєте перевірити, що `res.text` == `'hello Guest'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.text');
assert.match(data.assertions[1].args[1], /('|")hello Guest\1/);
```
