---
id: 587d824c367417b2b2512c4c
title: Твердження глибокої рівності за допомогою .deepEqual і .notDeepEqual
challengeType: 2
forumTopicId: 301587
dashedName: assert-deep-equality-with--deepequal-and--notdeepequal
---

# --description--

Нагадаємо, цей проєкт створюється на основі стартового проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`deepEqual()` перевіряє, чи є два об’єкти глибоко рівними.

# --instructions--

У межах `tests/1_unit-tests.js` під тестом з міткою `#7` у наборі `Equality` замініть кожен `assert` на `assert.deepEqual` або `assert.notDeepEqual`, щоб тест пройшов (повинен оцінюватися як `true`). Не змінюйте аргументи, передані у твердження.

# --hints--

Всі тести мають пройти.

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
})();
```

Ви маєте вибрати правильний метод для першого твердження — `deepEqual` чи `notDeepEqual`.

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'deepEqual',
    'The order of the keys does not matter'
  );
})();
```

Ви маєте вибрати правильний метод для другого твердження — `deepEqual` чи `notDeepEqual`.

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'notDeepEqual',
    'The position of elements within an array does matter'
  );
})();
```
