---
id: 587d824e367417b2b2512c56
title: "Перевірка, чи є значення певним типом структури даних"
challengeType: 2
forumTopicId: 301601
dashedName: test-if-a-value-is-of-a-specific-data-structure-type
---

# --description--

Нагадаємо, цей проєкт створюється на основі стартового проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`#typeOf` перевіряє, що тип значення відповідає заданому рядку, визначеному `Object.prototype.toString`.

# --instructions--

У `tests/1_unit-tests.js` у тесті з міткою `#17` у наборі тестів `Objects` замініть кожен `assert` на `assert.typeOf` або `assert.notTypeOf`, щоб тест пройшов (повинен оцінюватися як `true`). Не змінюйте аргументи, передані твердженням.

# --hints--

Всі тести мають пройти.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Виберіть правильний метод для першого твердження — `typeOf` чи `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'typeOf',
  'myCar is typeOf Object'
);
```

Виберіть правильний метод для другого твердження — `typeOf` чи `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'typeOf',
  'Car.model is a String'
);
```

Виберіть правильний метод для третього твердження — `typeOf` чи `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[2].method,
  'notTypeOf',
  'Plane.wings is a Number (not a String)'
);
```

Виберіть правильний метод для четвертого твердження — `typeOf` чи `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'typeOf',
  'Plane.engines is an Array'
);
```

Виберіть правильний метод для п’ятого твердження — `typeOf` чи `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[4].method,
  'typeOf',
  'Car.wheels is a Number'
);
```
