---
id: 587d824c367417b2b2512c4e
title: "Перевірка, чи одне значення менше або принаймні таке саме, як інше"
challengeType: 2
forumTopicId: 301606
dashedName: test-if-one-value-is-below-or-at-least-as-large-as-another
---

# --description--

Нагадаємо, цей проєкт створюється на основі стартового проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

У `tests/1_unit-tests.js` під тестом з міткою `#9` у наборі `Comparisons` замініть кожен `assert` на `assert.isBelow` або `assert.isAtLeast`, щоб тест пройшов (повинен оцінюватися як `true`). Не змінюйте аргументи, передані твердженням.

# --hints--

Всі тести мають пройти.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Виберіть правильний метод для першого твердження — `isBelow` чи `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isAtLeast',
  '5 is at least (>=) 5'
);
```

Виберіть правильний метод для другого твердження — `isBelow` чи `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'isAtLeast',
  '2 * Math.random() is at least 0'
);
```

Виберіть правильний метод для третього твердження — `isBelow` чи `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'isBelow', '1 is smaller than 2');
```

Виберіть правильний метод для четвертого твердження — `isBelow` чи `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'isBelow',
  '2/3 (0.6666) is smaller than 1'
);
```
