---
id: 587d824f367417b2b2512c5b
title: "Запустіть функціональні тести відповіді API за допомогою Chai-HTTP IV — метод PUT"
challengeType: 2
forumTopicId: 301591
dashedName: run-functional-tests-on-an-api-response-using-chai-http-iv---put-method
---

# --description--

Нагадаємо, цей проєкт створюється на основі стартового проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Ця вправа схожа на попередню.

Тепер, коли ви знаєте, як тестувати `PUT` запит, ваша черга зробити це з нуля.

# --instructions--

У межах `tests/2_functional-tests.js` змініть тест `'Send {surname: "da Verrazzano"}'` (`// #4`) і використайте методи `put` та `send` для тестування кінцевої точки `'/travellers'`.

Надішліть наступний JSON-об’єкт із вашим PUT-запитом:

```json
{
  "surname": "da Verrazzano"
}
```

Перевірте наступне у функції зворотного виклику `request.end`:

1.  `status` має бути `200`
2.  `type` має бути `application/json`
3.  `body.name` має бути `Giovanni`
4.  `body.surname` має бути `da Verrazzano`

Дотримуйтеся порядку тверджень вище — ми на нього покладаємося. Також обов’язково видаліть `assert.fail()` після завершення.

# --hints--

Усі тести мають пройти

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Ви маєте перевірити, що `res.status` дорівнює 200

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

Ви маєте перевірити, що `res.type` дорівнює `'application/json'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.type');
assert.match(data.assertions[1].args[1], /('|")application\/json\1/);
```

Ви маєте перевірити, що `res.body.name` дорівнює `'Giovanni'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'equal');
assert.equal(data.assertions[2].args[0], 'res.body.name');
assert.match(data.assertions[2].args[1], /('|")Giovanni\1/);
```

Ви маєте перевірити, що `res.body.surname` дорівнює `'da Verrazzano'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'equal');
assert.equal(data.assertions[3].args[0], 'res.body.surname');
assert.match(data.assertions[3].args[1], /('|")da Verrazzano\1/);
```
