---
id: 587d824b367417b2b2512c4b
title: Використовуйте потрійне рівність для перевірки суворої рівності
challengeType: 2
forumTopicId: 301610
dashedName: use-the-triple-equals-to-assert-strict-equality
---

# --description--

Нагадаємо, цей проєкт створюється на основі стартового проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`strictEqual()` порівнює об’єкти, використовуючи `===`.

# --instructions--

У `tests/1_unit-tests.js` у тесті з міткою `#6` у наборі `Equality` замініть кожен `assert` на `assert.strictEqual` або `assert.notStrictEqual`, щоб тест пройшов (повинен оцінюватися як `true`). Не змінюйте аргументи, передані твердженням.

# --hints--

Всі тести мають пройти.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Ви повинні вибрати правильний метод для першого твердження — `strictEqual` чи `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'notStrictEqual',
    'with strictEqual the type must match'
  );
```

Ви повинні вибрати правильний метод для другого твердження — `strictEqual` чи `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'strictEqual', '3*2 = 6...');
```

Ви повинні вибрати правильний метод для третього твердження — `strictEqual` чи `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'strictEqual',
    "6 * '2' is 12. Types match !"
  );
```

Ви повинні вибрати правильний метод для четвертого твердження — `strictEqual` чи `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[3].method,
    'notStrictEqual',
    'Even if they have the same elements, the Arrays are notStrictEqual'
  );
```
