---
id: 587d824e367417b2b2512c57
title: "Перевірка, чи є об’єкт екземпляром конструктора"
challengeType: 2
forumTopicId: 301605
dashedName: test-if-an-object-is-an-instance-of-a-constructor
---

# --description--

Нагадаємо, цей проєкт створюється на основі стартового проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`#instanceOf` перевіряє, чи є об’єкт екземпляром конструктора.

# --instructions--

У межах `tests/1_unit-tests.js` під тестом з міткою `#18` у наборі `Objects` змініть кожен `assert` на `assert.instanceOf` або `assert.notInstanceOf`, щоб тест пройшов (повинен оцінюватися як `true`). Не змінюйте аргументи, передані твердженням.

# --hints--

Всі тести мають пройти.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Ви маєте обрати правильний метод для першого твердження — `instanceOf` чи `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notInstanceOf',
  'myCar is not an instance of Plane'
);
```

Ви маєте обрати правильний метод для другого твердження — `instanceOf` чи `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'instanceOf',
  'airlinePlane is an instance of Plane'
);
```

Ви маєте обрати правильний метод для третього твердження — `instanceOf` чи `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[2].method,
  'instanceOf',
  'everything is an Object in JavaScript...'
);
```

Ви маєте обрати правильний метод для четвертого твердження — `instanceOf` чи `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'notInstanceOf',
  'myCar.wheels is not an instance of String'
);
```
