---
id: 587d824f367417b2b2512c5a
title: "Запустіть функціональні тести відповіді API за допомогою Chai-HTTP III — метод PUT"
challengeType: 2
forumTopicId: 301590
dashedName: run-functional-tests-on-an-api-response-using-chai-http-iii---put-method
---

# --description--

Нагадаємо, цей проєкт створюється на основі стартового проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Коли ви тестуєте `PUT` запит, ви часто надсилаєте разом із ним дані. Дані, які ви додаєте до свого `PUT` запиту, називаються тілом запиту.

Щоб надіслати `PUT` запит і JSON-об’єкт до кінцевої точки `'/travellers'`, ви можете використати `chai-http` плагін з методами `put` і `send`:

```js
chai
  .request(server)
  .keepOpen()
  .put('/travellers')
  .send({
    "surname": [last name of a traveller of the past]
  })
  ...
```

І маршрут відповідає таким чином:

```json
{
  "name": [first name],
  "surname": [last name],
  "dates": [birth - death years]
}
```

Перегляньте код сервера для різних відповідей на кінцеву точку `'/travellers'`.

# --instructions--

У межах `tests/2_functional-tests.js` змініть тест `'Send {surname: "Colombo"}'` (`// #3`) і використайте методи `put` і `send`, щоб протестувати кінцеву точку `'/travellers'`.

Надішліть наступний JSON-об’єкт із вашим PUT-запитом:

```json
{
  "surname": "Colombo"
}
```

Перевірте наступне у функції зворотного виклику `request.end`:

1.  `status` має бути `200`
2.  `type` має бути `application/json`
3.  `body.name` має бути `Cristoforo`
4.  `body.surname` має бути `Colombo`

Дотримуйтеся порядку тверджень вище — ми на нього покладаємося. Також обов’язково видаліть `assert.fail()` після завершення.

# --hints--

Усі тести мають пройти.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Ви маєте перевірити, що `res.status` дорівнює 200.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

Ви маєте перевірити, що `res.type` дорівнює `'application/json'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.type');
assert.match(data.assertions[1].args[1], /('|")application\/json\1/);
```

Ви маєте перевірити, що `res.body.name` дорівнює `'Cristoforo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'equal');
assert.equal(data.assertions[2].args[0], 'res.body.name');
assert.match(data.assertions[2].args[1], /('|")Cristoforo\1/);
```

Ви маєте перевірити, що `res.body.surname` дорівнює `'Colombo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'equal');
assert.equal(data.assertions[3].args[0], 'res.body.surname');
assert.match(data.assertions[3].args[1], /('|")Colombo\1/);
```
