---
id: 587d824d367417b2b2512c50
title: "Перевірка, чи є значення масивом"
challengeType: 2
forumTopicId: 301600
dashedName: test-if-a-value-is-an-array
---

# --description--

Нагадаємо, цей проєкт створюється на основі стартового проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

У `tests/1_unit-tests.js` у тесті з міткою `#11` у наборі `Arrays` замініть кожен `assert` на `assert.isArray` або `assert.isNotArray`, щоб тест пройшов (повинен оцінюватися як `true`). Не змінюйте аргументи, передані твердженням.

# --hints--

Всі тести мають пройти.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=10');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Вам слід обрати правильний метод для першого твердження — `isArray` чи `isNotArray`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=10');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isArray',
  'String.prototype.split() returns an Array'
);
```

Вам слід обрати правильний метод для другого твердження — `isArray` чи `isNotArray`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=10');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'isNotArray',
  'Array.prototype.indexOf() returns a number'
);
```
