---
id: 587d824b367417b2b2512c4a
title: Використайте подвійну рівність для підтвердження рівності
challengeType: 2
forumTopicId: 301609
dashedName: use-the-double-equals-to-assert-equality
---

# --description--

As a reminder, this project is being built upon the following starter project cloned from <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`equal()` порівнює об’єкти за допомогою `==`.

# --instructions--

У межах `tests/1_unit-tests.js` під тестом з міткою `#5` в наборі `Equality` змініть кожний `assert` на `assert.equal` або `assert.notEqual`, щоб пройти тест (повинен дорівнювати `true`). Не змінюйте аргументи, передані до тверджень.

# --hints--

Усі тести повинні пройти.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Ви повинні обрати правильний метод для першого твердження: `equal` або `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'equal',
    'Numbers are coerced into strings with == '
  );
```

Ви повинні обрати правильний метод для другого твердження: `equal` або `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'notEqual',
    ' == compares object references'
  );
```

Ви повинні обрати правильний метод для третього твердження: `equal` або `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'equal',
    "6 * '2' is 12 ! It should be equal to '12'"
  );
```

Ви повинні обрати правильний метод для четвертого твердження: `equal` або `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[3].method, 'notEqual', "6 + '2' is '62'...");
```

