---
id: 587d824b367417b2b2512c49
title: Тест на істинність
challengeType: 2
forumTopicId: 301596
dashedName: test-for-truthiness
---

# --description--

Нагадаємо, цей проєкт створюється на основі стартового проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`isTrue()` перевірятиме булеве значення `true`, а `isNotTrue()` пройде, якщо отримати будь-що, крім булевого значення `true`.

```js
assert.isTrue(true, 'This will pass with the boolean value true');
assert.isTrue('true', 'This will NOT pass with the string value "true"');
assert.isTrue(1, 'This will NOT pass with the number value 1');
```

`isFalse()` і `isNotFalse()` також існують і поводяться подібно до своїх істинних аналогів, але шукають булеве значення `false`.

# --instructions--

У `tests/1_unit-tests.js` у тесті з міткою `#4` у наборі `Basic Assertions` змініть кожен `assert` на `assert.isTrue` або `assert.isNotTrue`, щоб тест пройшов (повинен оцінюватися як `true`). Не змінюйте аргументи, передані твердженням.

# --hints--

Всі тести мають пройти.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Вам слід обрати правильний метод для першого твердження — `isTrue` чи `isNotTrue`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isTrue', 'True is true');
```

Вам слід обрати правильний метод для другого твердження — `isTrue` чи `isNotTrue`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'isTrue',
    'Double negation of a truthy value is true'
  );
```

Вам слід обрати правильний метод для третього твердження — `isTrue` чи `isNotTrue`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'isNotTrue',
    'A truthy object is not true - neither is a false one'
  );
```
