---
id: 587d8250367417b2b2512c5d
title: Запустіть функціональні тести за допомогою безголового браузера
challengeType: 2
forumTopicId: 301595
dashedName: run-functional-tests-using-a-headless-browser
---

# --description--

Нагадаємо, цей проєкт створюється на основі стартового проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

На сторінці є форма для введення даних. Вона надсилає дані на кінцеву точку `PUT /travellers` у вигляді AJAX-запиту.

Коли запит успішно завершується, код клієнта додає `<div>` з інформацією з відповіді до DOM.

Ось приклад, як використовувати Zombie.js для взаємодії з формою:

```js
test('Submit the surname "Polo" in the HTML form', function (done) {
  browser.fill('surname', 'Polo').then(() => {
    browser.pressButton('submit', () => {
      browser.assert.success();
      browser.assert.text('span#name', 'Marco');
      browser.assert.text('span#surname', 'Polo');
      browser.assert.elements('span#dates', 1);
      done();
    });
  });
});
```

Спочатку метод `fill` об’єкта `browser` заповнює поле `surname` форми значенням `'Polo'`. `fill` повертає проміс, тому до нього приєднується `then`.

Усередині функції зворотного виклику `then` використовується метод `pressButton` об’єкта `browser` для виклику обробника події `submit` форми. Метод `pressButton` є асинхронним.

Потім, коли отримано відповідь на AJAX-запит, робиться кілька тверджень, які підтверджують:

1.  Статус відповіді — `200`
2.  Текст у елементі `<span id='name'></span>` відповідає `'Marco'`
3.  Текст у елементі `<span id='surname'></span>` відповідає `'Polo'`
4.  Існує `1` елементів `<span id='dates'></span>`

Нарешті, викликається функція зворотного виклику `done`, яка потрібна через асинхронність тесту.

# --instructions--

Усередині `tests/2_functional-tests.js`, у тесті `'Submit the surname "Colombo" in the HTML form'` (`// #5`), автоматизуйте наступне:

1.  Заповніть форму прізвищем `Colombo`
2.  Натисніть кнопку відправки

А у функції зворотного виклику `pressButton`:

1.  Переконайтеся, що статус OK `200`
2.  Переконайтеся, що текст всередині елемента `span#name` дорівнює `'Cristoforo'`
3.  Переконайтеся, що текст всередині елемента `span#surname` дорівнює `'Colombo'`
4.  Переконайтеся, що елемент(и) `span#dates` існують і їх кількість дорівнює `1`

Не забудьте видалити виклик `assert.fail()`.

# --hints--

Всі тести мають пройти.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Ви маєте переконатися, що запит безголового браузера пройшов успішно.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'browser.success');
```

Ви маєте переконатися, що текст всередині елемента `span#name` дорівнює `'Cristoforo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'browser.text');
assert.match(data.assertions[1].args[0], /('|")span#name\1/);
assert.match(data.assertions[1].args[1], /('|")Cristoforo\1/);
```

Ви маєте переконатися, що текст всередині елемента `span#surname` дорівнює `'Colombo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'browser.text');
assert.match(data.assertions[2].args[0], /('|")span#surname\1/);
assert.match(data.assertions[2].args[1], /('|")Colombo\1/);
```

Ви маєте переконатися, що елемент `span#dates` існує і його кількість дорівнює 1.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'browser.elements');
assert.match(data.assertions[3].args[0], /('|")span#dates\1/);
assert.equal(data.assertions[3].args[1], 1);
```
