---
id: 5f8884f4c46685731aabfc41
title: Запуск функціональних тестів за допомогою безголового браузера II
challengeType: 2
forumTopicId: 301594
dashedName: run-functional-tests-using-a-headless-browser-ii
---

# --description--

Нагадаємо, цей проєкт створюється на основі шаблонного коду, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

У межах `tests/2_functional-tests.js`, у тесті `'Submit the surname "Vespucci" in the HTML form'` (`// #6`) автоматизуйте наступне:

1. Заповніть форму прізвищем `Vespucci`
2. Натисніть кнопку відправки

А у функції зворотного виклику `pressButton`:

1. Перевірте, що статус є OK `200`
2. Перевірте, що текст всередині елемента `span#name` дорівнює `'Amerigo'`
3. Перевірте, що текст всередині елемента `span#surname` дорівнює `'Vespucci'`
4. Перевірте, що елемент(и) `span#dates` існують і їх кількість дорівнює `1`

Не забудьте видалити виклик `assert.fail()`.

# --hints--

Всі тести мають пройти.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Ви повинні перевірити, що запит безголового браузера пройшов успішно.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'browser.success');
```

Ви повинні перевірити, що текст всередині елемента `span#name` дорівнює `'Amerigo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'browser.text');
assert.match(data.assertions[1].args[0], /('|")span#name\1/);
assert.match(data.assertions[1].args[1], /('|")Amerigo\1/);
```

Ви повинні перевірити, що текст всередині елемента `span#surname` дорівнює `'Vespucci'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'browser.text');
assert.match(data.assertions[2].args[0], /('|")span#surname\1/);
assert.match(data.assertions[2].args[1], /('|")Vespucci\1/);
```

Ви повинні перевірити, що елемент `span#dates` існує і його кількість дорівнює 1.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'browser.elements');
assert.match(data.assertions[3].args[0], /('|")span#dates\1/);
assert.equal(data.assertions[3].args[1], 1);
```
