---
id: 68de40f53fcaf106f51cc20f
title: Construa um Reprodutor de Áudio e Vídeo HTML
challengeType: 25
dashedName: build-an-html-audio-and-video-player
demoType: onClick
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter um elemento `h1` para o título principal da página.
2. Você deve ter dois elementos `section`.
3. Dentro do primeiro elemento `section`, você deve ter um elemento `h2` para o título do vídeo que está sendo reproduzido.
4. Abaixo do elemento `h2`, você deve ter um elemento `video` com os atributos `controls` e `width`. O atributo `width` deve ser definido como `640`.
5. Dentro do elemento `video`, você deve ter um elemento `source` com um atributo `src` apontando para um arquivo de vídeo e um atributo `type`.
   - Você pode usar `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`.
6. Dentro do segundo elemento `section`, você deve ter um elemento `h2` para o título da música que está tocando.
7. Abaixo do elemento `h2`, você deve ter um elemento `audio` com os atributos `controls` e `loop` e um atributo `src` que aponta para um arquivo de áudio.
   - Você pode usar `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`.
   - Ou `https://cdn.freecodecamp.org/curriculum/js-music-player/we-are-going-to-make-it.mp3`.

# --hints--

Você deve ter um elemento `h1` para o título principal da página.

```js
const h1 = document.querySelector('h1');
assert.exists(h1);
```

Seu `h1` não deve estar vazio.

```js
const h1 = document.querySelector('h1');
assert.isNotEmpty(h1.textContent.trim());
```

Você deve ter apenas um elemento `h1`

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Você deve ter dois elementos `section`.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);

```

Dentro do primeiro elemento `section`, você deve ter um elemento `h2` para o título do vídeo que está sendo reproduzido.

```js
const section = document.querySelector('section');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

Abaixo do elemento `h2`, você deve ter um elemento `video`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
assert.exists(video);
```

O elemento `video` deve ter um atributo `controls`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const controlsAttr = video.controls;
assert.isTrue(controlsAttr);
```

O elemento `video` deve ter um atributo `width` definido como `640`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const widthAttr = video.getAttribute('width');
assert.equal(widthAttr, '640');
```

Dentro do elemento `video`, você deve ter um elemento `source`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
assert.exists(source);
```

O elemento `source` deve ter um atributo `src` apontando para um arquivo de vídeo.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const srcAttr = source.getAttribute('src');
assert.match(srcAttr, /\.(mp4|webm|ogg|avi|mov)$/i);
```

O elemento `source` deve ter um atributo `type`

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
assert.isNotEmpty(typeAttr);
```

O atributo `type` deve corresponder ao tipo de mídia da extensão do arquivo do atributo `src`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
const srcAttr = source.getAttribute('src');
const srcFileExt = srcAttr.match(/[^\.]+$/g)[0];
const fileExtToMediaTypes = {
  mp4: ['video/mp4'],
  webm: ['video/webm'],
  ogg: ['video/ogg'],
  avi: ['video/avi', 'video/vnd.avi', 'video/msvideo', 'video/x-msvideo'],
  mov: ['video/quicktime']
};
const srcMediaTypes = fileExtToMediaTypes[srcFileExt];
assert.oneOf(typeAttr, srcMediaTypes);
```

Dentro do segundo elemento `section`, você deve ter um elemento `h2` para o título da música que está tocando.

```js
const section = document.querySelector('section:nth-of-type(2)');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

Todos os seus elementos `h2` devem conter algum texto.

```js
const h2s = document.querySelectorAll('h2');
assert.isNotEmpty(h2s);
h2s.forEach((h2) => assert.isNotEmpty(h2.textContent.trim()));
```

Abaixo do elemento `h2`, você deve ter um elemento `audio`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
assert.exists(audio);
```

O elemento `audio` deve ter um atributo `controls`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const controlsAttr = audio.controls;
assert.isTrue(controlsAttr);
```

O elemento `audio` deve ter um atributo `loop`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const loopAttr = audio.loop;
assert.isTrue(loopAttr);
```

O elemento `audio` deve ter um atributo `src` apontando para um arquivo de áudio.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const srcAttr = audio.getAttribute('src');
assert.match(srcAttr, /\.(mp3|wav|ogg|aac|flac)$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>
  <main>
    <h1>HTML Audio and Video Lab</h1>

    <section>
      <h2>What is the map method and how does it work</h2>
      <video controls width="640">
        <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4" />
      </video>
    </section>

    <section>
      <h2>Sailing Away</h2>
      <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" controls loop>
      </audio>
    </section>

  </main>
</body>
</html>
```
