---
id: 587d7b8d367417b2b2512b5b
title: Aprender sobre programação funcional
challengeType: 1
forumTopicId: 301233
dashedName: learn-about-functional-programming
---

# --description--

Programação funcional é um estilo de programação onde as soluções são funções simples e isoladas, sem quaisquer efeitos colaterais fora do escopo da função: `ENTRADA -> PROCESSO -> SAÍDA`

Quando falamos de programação funcional, pensamos em:

Funções isoladas - não há dependência do estado do programa, o que inclui variáveis globais que estão sujeitas a alterações

Funções puras - a mesma entrada sempre produz a mesma saída

Funções com efeitos colaterais limitados - quaisquer alterações, ou mutações, no estado do programa fora da função são cuidadosamente controladas

# --instructions--

Os membros do freeCodeCamp adoram chá.

As funções `prepareTea` e `getTea` já estão definidas para você no editor de texto. Chame a função `getTea` para adquirir 40 xícaras de chá para o time e as armazene na variável `tea4TeamFCC`.

# --hints--

A variável `tea4TeamFCC` deve conter 40 xícaras de chá para o time.

```js
assert(tea4TeamFCC.length === 40);
```

A variável `tea4TeamFCC` deve conter xícaras de chá verde.

```js
assert(tea4TeamFCC[0] === 'greenTea');
```

# --seed--

## --seed-contents--

```js
// Function that returns a string representing a cup of green tea
const prepareTea = () => 'greenTea';

/*
Given a function (representing the tea type) and number of cups needed, the
following function returns an array of strings (each representing a cup of
a specific type of tea).
*/
const getTea = (numOfCups) => {
  const teaCups = [];

  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }
  return teaCups;
};

// Only change code below this line
const tea4TeamFCC = null;
// Only change code above this line
```

# --solutions--

```js
const prepareTea = () => 'greenTea';

const getTea = (numOfCups) => {
  const teaCups = [];
  
  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }

  return teaCups;
};

const tea4TeamFCC = getTea(40); 
```
