---
id: 67ef8f695db52583424bd118
title: Diseñar una página de reseña de películas
challengeType: 25
dashedName: design-a-movie-review-page
demoType: onClick
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Deberías tener un elemento `main`.
2. Dentro del elemento `main`, deberías tener un elemento `h1` para el título de la película.
3. Debajo del elemento `h1`, deberías tener un elemento `img` que muestre la portada de la película. El elemento `img` debe tener un texto `alt` descriptivo que describa la imagen. Eres libre de usar la siguiente imagen si lo deseas: `https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png`.
4. Deberías tener un elemento `p` que contenga una breve descripción de la película.
5. Deberías tener otro elemento `p` para mostrar la calificación de la película. Dentro de él, deberías tener estos elementos en el orden indicado:
   - Un elemento `strong` con el texto `Calificación de la película:`.
   - Un elemento `span` con un atributo `aria-hidden` establecido como `true`, que contenga una representación visual de la calificación usando estrellas `⭐⭐⭐⭐⭐⭐⭐⭐⭐☆`.
   - Un valor numérico, que representa la calificación, entre paréntesis (por ejemplo, `9.2/10`) después del span.
6. Deberías tener un elemento `h2` con el texto `Miembros del elenco`.
7. Debes tener un elemento `ul`.
8. Dentro del elemento `ul`, deberías tener múltiples elementos `li`, cada uno conteniendo un elemento `strong` para el nombre del actor seguido por el nombre del personaje correspondiente precedido por el texto `como`. (por ejemplo, `James Holloway como Ethan Carter`).

# --hints--

Deberías tener un elemento `main`.

```js
assert.isNotNull(document.querySelector("main"));
```

Dentro del elemento `main`, deberías tener un elemento `h1` que contenga el título de la película.

```js
const h1 = document.querySelector("main > h1");
assert.isAbove(h1?.innerText.trim().length, 0);
```

Debajo del elemento `h1`, deberías tener un elemento `img` con un atributo `alt` descriptivo.

```js
const img = document.querySelector("main > h1 + img");
assert.isNotNull(img);
assert.isNotEmpty(img.getAttribute("alt"));
```

Deberías tener un elemento `p` después de la imagen que contenga la descripción de la película.

```js
const description = document.querySelector("main img + p");
assert.isNotNull(description);
assert.isNotEmpty(description.innerText.trim());
```

Deberías tener otro elemento `p` que contenga la calificación.

```js
const paragraphs = document.querySelectorAll("main p");
assert.isAtLeast(paragraphs.length, 2);
```

Dentro del segundo elemento `p`, deberías tener un elemento `strong` con el texto `Calificación de la película:`.

```js
const strongEl = document.querySelector("main p:nth-of-type(2) strong");
assert.match(strongEl?.innerText.trim(), /^Movie Rating:?$/);
```

Debería haber un elemento `span` dentro del párrafo de calificación.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.isNotNull(spanEl);
```

El elemento `span` dentro del párrafo de calificación debería tener diez estrellas, ya sea llenas (⭐) o vacías (☆), seguidas de un valor numérico entre paréntesis después del `span`.

```js  
const spanEl = document.querySelector("main p:nth-of-type(2) span");
const fullRatingText = document.querySelector("main p:nth-of-type(2)")?.textContent.replace(/\s+/g, ' ').trim();

assert.match(spanEl?.innerText.trim(), /^[⭐☆]{10}$/);
assert.match(fullRatingText, /Movie Rating:?\s*[⭐☆]{10}\s*\(\s*\d+(?:\.\d+)?\s*\/\s*10\s*\)$/);
```

El elemento `span` dentro del párrafo de calificación debería tener un atributo `aria-hidden` establecido como `true`.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.strictEqual(spanEl?.getAttribute("aria-hidden"), "true");
```

Deberías tener un elemento `h2`.

```js
assert.exists(document.querySelector("h2"));
```

Tu `h2` debería venir después de tu segundo elemento `p`.

```js
const allParagraphs = document.querySelectorAll("p");
assert.isAtLeast(allParagraphs.length, 2);

const h2 = document.querySelector("h2");
const main = document.querySelector("main");

assert.exists(h2);
assert.exists(main);

const elements = Array.from(main.querySelectorAll("*"));
const lastPIndex = elements.indexOf(allParagraphs[1]);
const h2Index = elements.indexOf(h2);

assert.isAbove(h2Index, lastPIndex);
```

Tu `h2` debería tener el texto `Miembros del elenco`.

```js
const h2 = document.querySelector("h2");
assert.equal(h2?.textContent.trim(), "Cast Members");
```

Deberías tener un elemento `ul` después del `h2`.

```js
const ul = document.querySelector("main h2 + ul");
assert.isNotNull(ul);
```

El `ul` debería contener múltiples elementos `li`.

```js
const liItems = document.querySelectorAll("main ul li");
assert.isAbove(liItems.length, 1);
```

Cada `li` debería contener un elemento `strong` para el nombre del actor, seguido por el nombre del personaje correspondiente precedido por el texto `como`. (por ejemplo, `James Holloway como Ethan Carter`).

```js
const liItems = document.querySelectorAll("main ul li");  
assert.isNotEmpty(liItems);
liItems.forEach(li => {
  const strong = li.querySelector("strong");
  assert.isNotNull(strong);
  const fullText = li.textContent.trim();
  const actorName = strong.textContent.trim();
  assert.isTrue(fullText.startsWith(actorName));
  const characterPart = fullText.slice(actorName.length).trim();
  assert.match(characterPart, /^as\s.+/);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
 <html lang="en">
 <head>
     <meta charset="UTF-8">
     <meta name="viewport" content="width=device-width, initial-scale=1.0">
     <title>Movie Review</title>
 </head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Movie Review</title>
  </head>
  <body>
    <main class="review">
      <h1>Movie Review: Rise Beyond</h1>
      <img
        src="https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png"
        alt="The Rise Beyond movie poster, starring Ethan Carter. A portrait of a young boy standing on a stage with his back toward the audience, holding a guitar at his side, staring at a backdrop of a cityscape with the sun rising behind its tall buildings."
      />
      <p>
        Rise Beyond is an uplifting drama that tells the story of Ethan Carter,
        a young man from a small town who dreams of becoming a world-class
        musician. Facing financial struggles and self-doubt, Ethan embarks on a
        journey of resilience, guided by his mentor Professor Adams and
        childhood friend Lena Mitchell. Through setbacks, failures, and moments
        of triumph, Rise Beyond reminds us that our greatest potential lies
        beyond our fears.
      </p>
      <p>
        <strong>Movie Rating:</strong>
        <span aria-hidden="true">⭐⭐⭐⭐⭐⭐⭐⭐⭐☆</span> (9.2/10)
      </p>

      <h2>Cast Members</h2>
      <ul>
        <li><strong>James Holloway</strong> as Ethan Carter</li>
        <li><strong>Olivia Sterling</strong> as Lena Mitchell</li>
        <li><strong>William Lancaster</strong> as Professor Adams</li>
        <li>
          <strong>Maria Collins</strong> as Evelyn Carter (Ethan’s Mother)
        </li>
      </ul>
    </main>
  </body>
</html>
```
