---
id: 66f686b8ebdb982fa8e14330
title: Tengeneza programu ya kubadilisha tarehe
challengeType: 26
dashedName: lab-date-conversion
---

# --description--

Katika maabara hii, utaandika msimbo wa programu ya kubadilisha tarehe ambayo hubadilisha tarehe iliyotolewa kuwa miundo tofauti. Kwa mfano, tarehe ya sasa `Fri Sep 27 2024 16:04:43 GMT+0500 (Pakistan Standard Time)` itapangwa kwa njia mbili zifuatazo:

- `(MM/DD/YYYY): 9/27/2024`.
- `(Month Day, Year): September 27, 2024`.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuunda kigezo kinachoitwa `currentDate` na kuweka thamani yake kuwa tarehe na wakati wa sasa kwa kutumia kitu cha `Date`.
2. Unapaswa kuunda kigezo kinachoitwa `currentDateFormat` na kuweka thamani yake kuwa mfuatano wa herufi `Current Date and Time: [current date]`. Badilisha `[current date]` na matokeo ya kigezo cha `currentDate`.
3. Unapaswa kuandika thamani ya `currentDateFormat` kwenye konsoli.
4. Unapaswa kuunda kitendakazi kinachoitwa `formatDateMMDDYYYY` ambacho kinapokea kitu cha tarehe kama kigezo. Unaweza kumtaja kigezo hiki kwa jina lolote unalotaka.
5. Kitendakazi chako cha `formatDateMMDDYYYY` kinapaswa kubadilisha tarehe ya sasa kuwa muundo wa `month/day/year` na kurudisha mfuatano wa herufi `Formatted Date (MM/DD/YYYY): [month]/[day]/[year]`.
6. Unapaswa kuunda kitendakazi kinachoitwa `formatDateLong` ambacho kinapokea kitu cha tarehe kama kigezo.
7. Kitendakazi chako cha `formatDateLong` kinapaswa kubadilisha tarehe ya sasa kuwa muundo wa `Month Day, Year` na kurudisha mfuatano wa herufi `Formatted Date (Month Day, Year): [formatted date]`.

**Kumbuka:** Ili vipimo vipite, hakikisha unatumia `en-US` kama eneo la lugha wakati wa kupanga tarehe.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `currentDate` kinachoshikilia tarehe na wakati wa sasa kwa kutumia kitu cha `Date`.

```js
// excluded timestamps
const expectedDate = new Date();
const expectedDateOnly = new Date(expectedDate.getFullYear(), expectedDate.getMonth(), expectedDate.getDate());

const currentDateOnly = new Date(currentDate.getFullYear(), currentDate.getMonth(), currentDate.getDate());

assert.equal(currentDateOnly.toString(), expectedDateOnly.toString());
```

Unapaswa kuwa na kigezo kinachoitwa `currentDateFormat` kinachoshikilia tarehe ya sasa katika muundo wa `Current Date and Time: <ddd> <MMM> <dd> <yyyy> <HH>:<mm>:<ss> <TIMEZONE>`.

```js
const expectedDateString = `Current Date and Time: `;
assert.include(currentDateFormat, expectedDateString);

const currentTimestamp = new Date(currentDateFormat.replace(expectedDateString, '')).getTime();
const expectedTimestamp = new Date().getTime();
assert.approximately(currentTimestamp, expectedTimestamp, 2000);
```

Unapaswa kuandika thamani ya `currentDateFormat` kwenye konsoli.

```js
assert.match(code, /console\.log\(\s*currentDateFormat\s*\)/);
```

Unapaswa kuwa na kitendakazi `formatDateMMDDYYYY` 

```js
assert.isFunction(formatDateMMDDYYYY);
```

Kitendakazi cha `formatDateMMDDYYYY` kinapaswa kupokea kigezo kimoja.

```js
assert.lengthOf(formatDateMMDDYYYY, 1);
```

Wakati kitu cha tarehe kinashikilia `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`, kitendakazi cha `formatDateMMDDYYYY` kinapaswa kurudisha `Formatted Date (MM/DD/YYYY): 9/27/2024`.

```js
let dateString = "Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)";
let dateObject = new Date(dateString);

assert.equal(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/27/2024");

// prevent hardcoding

dateString = "Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)";
dateObject = new Date(dateString);

assert.strictEqual(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/28/2024");
```

Unapaswa kuwa na kitendakazi `formatDateLong` 

```js
assert.isFunction(formatDateLong);
```

Kitendakazi cha `formatDateLong` kinapaswa kupokea kigezo kimoja.

```js
assert.lengthOf(formatDateLong, 1);
```

Wakati kitu cha tarehe kinashikilia `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`, kitendakazi cha `formatDateLong` kinapaswa kurudisha `Formatted Date (Month Day, Year): September 27, 2024`.

```js
let testDate = new Date('Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)');
let result = formatDateLong(testDate);

assert.strictEqual(result,'Formatted Date (Month Day, Year): September 27, 2024', `Test failed: ${result}`); 

// prevent hardcoding

testDate = new Date('Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)');
result = formatDateLong(testDate);
assert.strictEqual(result, 'Formatted Date (Month Day, Year): September 28, 2024'); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const currentDate = new Date();
const currentDateFormat = `Current Date and Time: ${currentDate}`; 
console.log(currentDateFormat);

function formatDateMMDDYYYY(date) {
  const month = date.getMonth() + 1;
  const day = date.getDate();
  const year = date.getFullYear();
  return `Formatted Date (MM/DD/YYYY): ${month}/${day}/${year}`;
}

function formatDateLong(date) {
  const longDateFormat =  date.toLocaleDateString('en-US', {
    year: 'numeric',
    month: 'long',
    day: 'numeric'
  });
  return `Formatted Date (Month Day, Year): ${longDateFormat}`;
}

// usage
console.log(formatDateMMDDYYYY(currentDate));
console.log(formatDateLong(currentDate));

```
