---
id: 69446f18fdcb11aa516f7c5c
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Neste workshop, você criará uma função `countup`, que cria um array com números de 1 até um número dado usando recursão.

Comece criando uma função vazia chamada `countup` que recebe um único parâmetro.

# --hints--

Você deve ter uma função chamada `countup`.

```js
assert.isFunction(countup);
```

Sua função `countup` deve aceitar um único parâmetro.

```js
assert.strictEqual(countup.length, 1);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
