---
id: 69447054012a22f38c944ff0
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Adicione uma declaração `if` dentro da sua função `countup`.

Verifique se `number` é menor que `1` e retorne um array vazio.

Isto servirá como o caso base para sua recursão.

# --hints--

Quando `number` for menor que `1` a função deve retornar um array vazio.

```js
assert.deepStrictEqual(countup(0), []);
```

Quando `number` não for menor que `1` a função não deve retornar nada.

```js
assert.isUndefined(countup(1)); 
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
