---
id: 694470b9a65b278f981ec62e
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Finalmente, teste sua função `countup` chamando-a com o argumento `5` dentro de uma instrução `console.log()`. 

A saída deve ser `[1, 2, 3, 4, 5]`.

# --hints--

Você deve chamar a função `countup` com o argumento `5` dentro de uma instrução `console.log()`

```js
assert.deepStrictEqual(countup(5), [1, 2, 3, 4, 5]);
assert.match(code, /console\.log\s*\(\s*countup\s*\(\s*5\s*\)\s*\)/)
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
  } else {
    countArray = countup(number - 1);
    countArray.push(number);
    return countArray;
  }
}
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
  } else {
    countArray = countup(number - 1);
    countArray.push(number);
    return countArray;
  }
}
console.log(countup(5));
```
