---
id: 694870bcf3c6874112b6f951
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Para todos os números maiores ou iguais a `1`, crie um bloco `else` e dentro dele, chame a função `countup` com `number - 1` e atribua o resultado a `countArray`. 

Então use `.push()` para adicionar `number` a `countArray` e finalmente retorne `countArray`. 

Isso garante que a recursão eventualmente alcance o caso base e termine.

# --hints--

Você deve atribuir o resultado da chamada de `countup` com `number - 1` para `countArray`.

```js
assert.match(countup.toString(), /countArray\s*=\s*countup\(\s*number\s*-\s*1\s*\)/);
```

Você deve inserir `number` em `countArray`.

```js
assert.match(countup.toString(), /countArray\.push\(\s*number\s*\)/);
```

Você deve retornar o `countArray`.

```js
assert.deepStrictEqual(countup(3), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
