---
id: 6948705160b66cf93a916e29
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Declare um array vazio chamado `countArray` dentro da função `countup` logo acima da sua declaração `if`.

Este array armazenará o valor de `number` a cada chamada de função.

# --hints--

Você deve declarar um array vazio chamado `countArray`.

```js
assert.match(countup.toString(), /(let|const|var)\s+countArray\s*=\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region-- 
  
--fcc-editable-region--
  if (number < 1) {
    return [];
  }
}
```
