---
id: 66f17db06803d11a1bd19a20
title: Ein Quizspiel erstellen
challengeType: 26
dashedName: lab-quiz-game
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten ein Array mit dem Namen `questions` erstellen.
1. Das `questions`-Array sollte mindestens fünf Objekte enthalten, die jeweils die Schlüssel `category`, `question`, `choices` und `answer` besitzen.
1. Der Schlüssel `category` sollte den Wert eines Strings haben, der eine Fragekategorie darstellt.
1. Der Schlüssel `question` sollte den Wert eines Strings haben, der eine Frage darstellt.
1. Der Schlüssel `choices` sollte den Wert eines Arrays enthalten, das drei Strings mit alternativen Antworten auf die Frage enthält.
1. Der Schlüssel `answer` sollte den Wert eines Strings haben, der die richtige Antwort auf die Frage darstellt. Außerdem sollte der Wert von `answer` im `choices`-Array enthalten sein.
1. Sie sollten eine Funktion mit dem Namen `getRandomQuestion` haben, die ein Array von Fragen als Parameter nimmt und ein zufälliges Frageobjekt aus dem Array zurückgibt.
1. Sie sollten eine Funktion mit dem Namen `getRandomComputerChoice` haben, die das Array der verfügbaren Auswahlmöglichkeiten als Parameter nimmt und eine zufällige Antwort auf die ausgewählte Frage zurückgibt.
1. Sie sollten eine Funktion mit dem Namen `getResults` haben, die das Frageobjekt als ersten Parameter und die Wahl des Computers als zweiten Parameter nimmt. Die Funktion sollte `The computer's choice is correct!` zurückgeben, wenn die Antwort korrekt ist. Andernfalls gibt sie `The computer's choice is wrong. The correct answer is: <correct-answer>` zurück, wobei `<correct-answer>` der Wert der richtigen Antwort auf die gewählte Frage ist.

# --hints--

Sie sollten ein Array mit dem Namen `questions` erstellen.

```js
assert.isArray(questions);
```

Das `questions`-Array sollte mindestens fünf Objekte enthalten, die jeweils die Schlüssel `category`, `question`, `choices` und `answer` besitzen.

```js
assert.isAtLeast(questions.length, 5);
questions.forEach(question => {
    assert.isObject(question);
    assert.hasAllKeys(question, ['category', 'choices', 'question', 'answer']);
});

```

Der Schlüssel `category` sollte den Wert eines Strings haben, der eine Fragekategorie darstellt.

```js
assert.isNotEmpty(questions);
questions.forEach(({category}) => {
    assert.isString(category);
    assert.isNotEmpty(category);
});
```

Der Schlüssel `question` sollte den Wert eines Strings haben, der eine Frage darstellt.

```js
assert.isNotEmpty(questions);
questions.forEach(({question}) => {
    assert.isString(question);
    assert.include(question, '?');
    assert.isNotEmpty(question.replaceAll('?', ''));
})
```

Der Schlüssel `choices` sollte den Wert eines Arrays enthalten, das drei unterschiedliche Strings enthält.

```js
assert.isNotEmpty(questions);
questions.forEach(({choices}) => {
    assert.isArray(choices);
    assert.lengthOf(choices, 3);
    assert.lengthOf(new Set(choices), 3);
    choices.forEach(choice => {
        assert.isString(choice);
        assert.isNotEmpty(choice);
    });
})
```

Der Schlüssel `answer` sollte den Wert eines Strings haben.

```js
assert.isNotEmpty(questions);
questions.forEach(({answer}) => {
    assert.isString(answer);
    assert.isNotEmpty(answer);
});
```

Der Wert von `answer` sollte im `choices`-Array enthalten sein.

```js
assert.isNotEmpty(questions);
questions.forEach(({answer, choices}) => assert.oneOf(answer, choices));
```

Sie sollten eine Funktion mit dem Namen `getRandomQuestion` haben, die ein Array von Fragen als Parameter nimmt und ein zufälliges Frageobjekt aus dem Array zurückgibt.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getRandomQuestion);
const parameters = explorer.allFunctions.getRandomQuestion.parameters;
assert.lengthOf(parameters, 1);
assert.isObject(getRandomQuestion(questions));
assert.deepInclude(questions, getRandomQuestion(questions), "getRandomQuestion did not return one of the objects inside questions");
```

Sie sollten eine Funktion mit dem Namen `getRandomComputerChoice` haben, die das Array der verfügbaren Auswahlmöglichkeiten als Parameter nimmt und eine zufällige Antwort auf die ausgewählte Frage zurückgibt.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getRandomComputerChoice);
const randomArray = ['a', 'b', 'c'];
assert.oneOf(getRandomComputerChoice(randomArray), randomArray)
```

Sie sollten eine Funktion mit dem Namen `getResults` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getResults);
```

Ihre `getResults`-Funktion sollte das Frageobjekt als ersten Parameter und die Wahl des Computers als zweiten Parameter nehmen.

```js
const testQuestion = questions[0];
const wrongChoice = questions[0].choices.find(choice => choice !== questions[0].answer);
assert.equal(getResults(testQuestion, testQuestion.answer), `The computer's choice is correct!`);
assert.equal(getResults(testQuestion, wrongChoice), `The computer's choice is wrong. The correct answer is: ${testQuestion.answer}`);
```

Wenn die Wahl des Computers mit der Antwort übereinstimmt, sollte `getResults` `The computer's choice is correct!` zurückgeben.

```js
assert.equal(getResults(questions[0], questions[0].answer), `The computer's choice is correct!`);
```

Wenn die Wahl des Computers nicht mit der Antwort übereinstimmt, sollte `getResults` `The computer's choice is wrong. The correct answer is: <correct-answer>` zurückgeben, wobei `<correct-answer>` der Wert der richtigen Antwort auf die gewählte Frage ist.

```js
assert.equal(getResults({category: 'misc', choices: ['a', 'b', 'c'], question: "question?", answer: "b"}, "a"), `The computer's choice is wrong. The correct answer is: b`)
```

Ihre `getResults`-Funktion sollte eine exakte Gleichheitsprüfung verwenden, keine Teilstringübereinstimmung.

```js
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Hamburger"}, "Ham"), `The computer's choice is wrong. The correct answer is: Hamburger`);
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Ham"}, "Hamburger"), `The computer's choice is wrong. The correct answer is: Ham`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const questions = [
    {
        category: "science",
        question: "What is the chemical symbol for potassium?",
        choices: ["P", "K", "Pt"],
        answer: "K"
    },
    {
        category: "science",
        question: "What is the unit of electrical resistance?",
        choices: ["Ohm", "Coulomb", "Sievert"],
        answer: "Ohm"
    },
    {
        category: "geography",
        question: "What is the capital city of Australia?",
        choices: ["Sidney", "Canberra", "Wellington"],
        answer: "Canberra"
    },
    {
        category: "literature",
        question: 'Who wrote "1984"?',
        choices: ["Ray Bradbury", "Aldous Huxley", "George Orwell"],
        answer: "George Orwell"
    },
    {
        category: "sport",
        question: "How many players are on a standard volleyball team?",
        choices: ["6", "7", "12"],
        answer: "6"
    }
]

function getRandomQuestion(questionList) {
    const randomIndex = Math.floor(Math.random() * questionList.length);
    return questionList[randomIndex];
}

function getRandomComputerChoice(choices) {
    const randomIndex = Math.floor(Math.random() * choices.length);
    return choices[randomIndex];
}

function getResults(obj, computerChoice) {
    return computerChoice === obj.answer
        ? "The computer's choice is correct!"
        : `The computer's choice is wrong. The correct answer is: ${obj.answer}`;
}

const questionObj = getRandomQuestion(questions);
const { question, choices } = questionObj;

console.log(question);
console.log(`Choices: ${choices}`);

const computerChoice = getRandomComputerChoice(choices);
console.log(`Computer chooses: ${computerChoice}`);

const results = getResults(questionObj, computerChoice);
console.log(results);

```
