---
id: 5d822fd413a79914d39e98d0
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Erstellen Sie ein `div`-Element im `body` mit einer Klasse von `background-buildings`. Dies wird ein Container für eine Gruppe von Gebäuden sein.

# --hints--

Sie sollten ein `div`-Element erstellen.

```js
assert.exists(document.querySelector('div'));
```

Ihr `div`-Element sollte sich innerhalb des `body` befinden.

```js
assert.equal(document.querySelector('div')?.parentElement?.localName, 'body');
```

Ihr `div`-Element sollte eine Klasse von `background-buildings` haben.

```js
assert.isTrue([...document.querySelector('div')?.classList]?.includes('background-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

```

