---
id: 5d822fd413a79914d39e98cc
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Willkommen im CSS-Variablen-Skyline-Projekt! Start by adding a `link` element that links your `styles.css` file within the `head` element.

# --hints--

Dein Code sollte ein `link`-Element enthalten.

```js
assert.isNotNull(document.querySelector('link'));
```

Du solltest nicht die vorhandenen `head`-Tags ändern. Stelle sicher, dass du das schließende Tag nicht gelöscht hast.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

You should have one void `link` element.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

Dein `link`-Element sollte sich innerhalb deines `head`-Elements befinden.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Dein `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` enthalten.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Dein `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` enthalten.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
