---
id: 5d822fd413a79914d39e98cf
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Sie sehen das `body` (es ist das innerste Kästchen auf Ihrer Seite); das Kästchen darum herum ist das `html`-Element. Lassen Sie Ihr `body` den gesamten Viewport ausfüllen, indem Sie ihm eine `height` von `100vh` geben. Entfernen Sie den Standard-`margin` vom `body`, indem Sie die `margin` auf `0` setzen. Stellen Sie schließlich die `overflow`-Eigenschaft auf `hidden` ein, um alle Scrollleisten auszublenden, die erscheinen, wenn etwas über den Viewport hinausgeht.

# --hints--

Sie sollten den `body`-Selektor verwenden.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Ihr `body` sollte eine `height` von `100vh` haben.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

Ihr `body` sollte eine `margin` von `0` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

Ihr `body` sollte die `overflow`-Eigenschaft auf `hidden` gesetzt haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--
    
```

