---
id: 5d822fd413a79914d39e98cc
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Willkommen beim CSS-Variablen-Skyline-Projekt! Beginnen Sie damit, ein `link`-Element hinzuzufügen, das Ihre `styles.css`-Datei innerhalb des `head`-Elements verlinkt.

# --hints--

Ihr Code sollte ein `link`-Element enthalten.

```js
assert.isNotNull(document.querySelector('link'));
```

Sie sollten Ihre vorhandenen `head`-Tags nicht ändern. Stellen Sie sicher, dass Sie das schließende Tag nicht gelöscht haben.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

Sie sollten ein leeres `link`-Element haben.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

Ihr `link`-Element sollte sich innerhalb Ihres `head`-Elements befinden.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Ihr `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Ihr `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>    
--fcc-editable-region--
    
--fcc-editable-region--
  </head>

  <body>
  </body>
</html>
```
