---
id: 683ecb010f4c6bafd428802c
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Workshop validieren Sie einen Satz medizinischer Daten, um sicherzustellen, dass sie einer Reihe von Regeln entsprechen.

Deklarieren Sie eine Variable namens `medical_records` und weisen Sie ihr eine leere Liste zu. In den folgenden Schritten werden Sie diese verwenden, um Ihre medizinischen Daten zu speichern.

# --hints--

Sie sollten eine Variable namens `medical_records` deklarieren.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable("medical_records")
`) })
```

Sie sollten Ihrer Variable `medical_records` eine leere Liste zuweisen.

```js
({ test: () => runPython(`
assert medical_records == []
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

