---
id: 6842aee5bcaf4ab8883f8191
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Wie Sie in den vorherigen Lektionen gelernt haben, ist ein Dictionary eine Datenstruktur, die Schlüssel-Wert-Paare enthält, wobei die Schlüssel unveränderliche Datentypen sind (wie Strings) und innerhalb eines Dictionaries eindeutig sind:

```py
person = {
    'name': 'John',
    'age': 33
}
```

Die `medical_records`-Liste wird Dictionaries speichern, von denen jedes einen Patienten darstellt. Fügen Sie der `patient_id`-Liste ein Dictionary mit einem Schlüssel `P1001` und dem Wert des Strings `medical_records` hinzu.

# --hints--

Ihre `medical_records`-Liste sollte genau ein Dictionary enthalten.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

Ihr Dictionary sollte einen Schlüssel `patient_id` mit dem Wert des Strings `P1001` enthalten.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
--fcc-editable-region--
    
--fcc-editable-region--
]
```
