---
id: 6842d659a4ae7973dabab3b0
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Vervollständigen Sie das Dictionary, indem Sie die folgenden drei Schlüssel-Wert-Paare hinzufügen:

- Einen Schlüssel `diagnosis` mit dem Wert `Hypertension`
- Einen Schlüssel `medications` mit dem Wert der Liste `['Lisinopril']`
- Einen Schlüssel `last_visit_id` mit dem Wert `V2301`.

# --hints--

Sie sollten in Ihrem Dictionary einen Schlüssel `diagnosis` mit dem Wert des Strings `Hypertension` haben.

```js
({ test: () => runPython(`
assert medical_records[0].get("diagnosis", None) == "Hypertension"
`) })
```

Sie sollten in Ihrem Dictionary einen Schlüssel `medications` mit dem Wert der Liste `['Lisinopril']` haben.

```js
({ test: () => runPython(`
assert medical_records[0].get("medications", None) == ['Lisinopril']
`) })
```

Sie sollten in Ihrem Dictionary einen Schlüssel `last_visit_id` mit dem Wert des Strings `V2301` haben.

```js
({ test: () => runPython(`
assert medical_records[0].get("last_visit_id", None) == "V2301"
`) })
```

Sie sollten in Ihrem Dictionary einen Schlüssel `patient_id` mit dem Wert des Strings `P1001` haben.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

Sie sollten in Ihrem Dictionary einen Schlüssel `age` mit dem Wert der Ganzzahl `34` haben.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

Sie sollten in Ihrem Dictionary einen Schlüssel `gender` mit dem Wert des Strings `Female` haben.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
--fcc-editable-region--
        'gender': 'Female'
        
--fcc-editable-region--
    }    
]
```
