---
id: 6842d5b537c64266420faf62
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Fügen Sie Ihrem Wörterbuch einen Schlüssel `age` mit dem Wert der Ganzzahl `34` und einen Schlüssel `gender` mit dem Wert des Strings `Female` hinzu. Vergessen Sie nicht das Komma zwischen den Schlüssel-Wert-Paaren.

# --hints--

Ihre `medical_records`-Liste sollte ein einzelnes Wörterbuch enthalten.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

Ihr Wörterbuch sollte einen Schlüssel `age` mit dem Wert der Ganzzahl `34` enthalten.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

Ihr Wörterbuch sollte einen Schlüssel `gender` mit dem Wert des Strings `Female` enthalten.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

Ihr Wörterbuch sollte einen Schlüssel `patient_id` mit dem Wert des Strings `P1001` enthalten.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
--fcc-editable-region--
        'patient_id': 'P1001'
        
--fcc-editable-region--
    }    
]
```
