---
id: 6842ba0086201089ea268c92
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Entsprechend der Struktur, die Sie in den vorherigen Schritten verwendet haben, wurde die `medical_records`-Liste für Sie mit den Daten anderer Patientinnen und Patienten gefüllt. Schauen Sie sie sich gerne an.

Als Nächstes beginnen Sie mit dem Schreiben der Funktion zur Validierung des Datensatzes. Erstellen Sie eine Funktion namens `validate` mit einem einzelnen Parameter `data`.

Sie möchten sicherstellen, dass Ihre Daten entweder eine Liste oder ein Tuple sind. Deklarieren Sie daher innerhalb der `validate`-Funktion eine Variable namens `is_sequence` und weisen Sie ihr einen Aufruf von `isinstance` zu. Übergeben Sie `data` als erstes Argument und ein Tuple, das `list` und `tuple` enthält, als zweites Argument.

# --hints--

Sie sollten eine Funktion namens `validate` haben.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("validate")
`) })
```

Ihre `validate`-Funktion sollte einen einzelnen Parameter `data` haben.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_args("data")
`) })
```

Sie sollten innerhalb Ihrer `is_sequence`-Funktion eine Variable namens `validate` haben.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_variable("is_sequence")
`) })
```

Sie sollten `isinstance(data, (list, tuple))` Ihrer `is_sequence`-Variablen zuweisen.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("validate").find_variable("is_sequence")
assert _cond.is_equivalent("is_sequence = isinstance(data, (list, tuple))") or _cond.is_equivalent("is_sequence = isinstance(data, (tuple, list))")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

--fcc-editable-region--

--fcc-editable-region--
```
