---
id: 6846a1f2805ac73ce671a5d8
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Wie Sie in einer vorherigen Lektion gelernt haben, ermöglicht Ihnen die `enumerate`-Funktion, den Index beim Durchlaufen eines iterierbaren Objekts zu verfolgen:

```py
person = {'name': 'John', 'age': 33}

for index, item in enumerate(person):
    print(index, item)

# 0 name
# 1 age
```

Erstellen Sie eine `for`-Schleife, die `data` durchläuft. Verwenden Sie die `enumerate`-Funktion, um bei jeder Iteration sowohl den Index als auch das Element in `data` zu erhalten. Verwenden Sie `index` und `dictionary` als Iterationsvariablen.

Verwenden Sie vorerst `pass`, um den Schleifenkörper zu füllen.

# --hints--

Sie sollten eine `for`-Schleife in Ihrer `validate`-Funktion haben.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0]
`) })
```

Ihre `for`-Schleife sollte `enumerate(data)` durchlaufen.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_for_iter().is_equivalent("enumerate(data)")
`) })
```

Sie sollten `index` und `dictionary` als Iterationsvariablen in Ihrer `for`-Schleife verwenden.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_for_vars().is_equivalent("index, dictionary")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False

--fcc-editable-region--
    
--fcc-editable-region--
```
