---
id: 6849311be42a0431b8e59109
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Sie prüfen, ob die an Ihre Funktion übergebenen Daten eine Liste oder ein Tuple sind. Sie müssen außerdem sicherstellen, dass jedes Element in der Folge ein Wörterbuch ist.

Innerhalb Ihrer `for`-Schleife, wenn das Element in `dictionary` keine Instanz von `dict` ist, geben Sie `Invalid format: expected a dictionary at position <index>.` aus (wobei `<index>` durch den aktuellen Index ersetzt werden sollte) und setzen `is_invalid` auf `True`.

# --hints--

Sie sollten eine `if`-Anweisung innerhalb Ihrer `for`-Schleife haben.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[0]
`) })
```

Ihre `if`-Anweisung sollte prüfen, ob `dictionary` keine Instanz von `dict` ist.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(dictionary, dict)")
`) })
```

Sie sollten `Invalid format: expected a dictionary at position <index>.` ausgeben (wobei `<index>` durch den aktuellen Index ersetzt werden sollte) innerhalb Ihrer `if`-Anweisung.

```js
({ test: () => assert(runPython(`
built_in_print = print
out = []

def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
validate([{}, "", {}, [], 0])
print = built_in_print

_first = 'Invalid format: expected a dictionary at position 1.'
_second = 'Invalid format: expected a dictionary at position 3.'
_third = 'Invalid format: expected a dictionary at position 4.'
len(out) >= 3 and _first in out and _second in out and _third in out
`)) })
```

Sie sollten `is_invalid` auf `True` setzen innerhalb Ihrer `if`-Anweisung.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("is_invalid = True")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False

    for index, dictionary in enumerate(data):
--fcc-editable-region--
        pass
--fcc-editable-region--
```
