---
id: 66f17db06803d11a1bd19a20
title: 建置一個測驗遊戲
challengeType: 26
dashedName: lab-quiz-game
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該創建一個名為 `questions` 的陣列。
1. `questions` 陣列應該至少包含五個物件，每個物件都具有 `category`、`question`、`choices` 和 `answer` 這些鍵。
1. `category` 鍵應該有一個表述問題類別的字串值。
1. `question` 鍵應該有一個表述問題的字串值。
1. `choices` 鍵的值應該是一個包含三個字串的陣列，這些字串是該問題的替代答案。
1. `answer` 鍵應該有一個字串值，表述問題的正確答案。此外，`answer` 的值應該包含在 `choices` 陣列中。
1. 你應該有一個名為 `getRandomQuestion` 的函式（程式），它接受一個問題陣列作為參數，並從該陣列中傳回一個隨機的問題物件。
1. 你應該有一個名為 `getRandomComputerChoice` 的函式（程式），它以可用選項的陣列作為參數，並傳回對所選問題的隨機答案。
1. 你應該有一個名為 `getResults` 的函式（程式），它以問題物件作為第一個參數，以電腦的選擇作為第二個參數。如果答案正確，該函式應傳回 `The computer's choice is correct!`。否則，它會傳回 `The computer's choice is wrong. The correct answer is: <correct-answer>`，其中 `<correct-answer>` 是所選問題的正確答案值。

# --hints--

你應該創建一個名為 `questions` 的陣列。

```js
assert.isArray(questions);
```

`questions` 陣列應該至少包含五個物件，每個物件都具有 `category`、`question`、`choices` 和 `answer` 這些鍵。

```js
assert.isAtLeast(questions.length, 5);
questions.forEach(question => {
    assert.isObject(question);
    assert.hasAllKeys(question, ['category', 'choices', 'question', 'answer']);
});

```

`category` 鍵應該有一個表述問題類別的字串值。

```js
assert.isNotEmpty(questions);
questions.forEach(({category}) => {
    assert.isString(category);
    assert.isNotEmpty(category);
});
```

`question` 鍵應該有一個表述問題的字串值。

```js
assert.isNotEmpty(questions);
questions.forEach(({question}) => {
    assert.isString(question);
    assert.include(question, '?');
    assert.isNotEmpty(question.replaceAll('?', ''));
})
```

`choices` 鍵的值應該是一個包含三個彼此不同字串的陣列。

```js
assert.isNotEmpty(questions);
questions.forEach(({choices}) => {
    assert.isArray(choices);
    assert.lengthOf(choices, 3);
    assert.lengthOf(new Set(choices), 3);
    choices.forEach(choice => {
        assert.isString(choice);
        assert.isNotEmpty(choice);
    });
})
```

`answer` 鍵應該具有字串的值。

```js
assert.isNotEmpty(questions);
questions.forEach(({answer}) => {
    assert.isString(answer);
    assert.isNotEmpty(answer);
});
```

`answer` 的值應該包含在 `choices` 陣列中。

```js
assert.isNotEmpty(questions);
questions.forEach(({answer, choices}) => assert.oneOf(answer, choices));
```

你應該有一個名為 `getRandomQuestion` 的函式（程式），它接受一個問題陣列作為參數，並從該陣列中傳回一個隨機的問題物件。

```js
const parameters = __helpers.getFunctionParams(getRandomQuestion.toString());

assert.isFunction(getRandomQuestion);
assert.lengthOf(parameters, 1);
assert.lengthOf(getRandomQuestion, 1);
assert.isObject(getRandomQuestion(questions));
assert.deepInclude(questions, getRandomQuestion(questions), "getRandomQuestion did not return one of the objects inside questions");
```

你應該有一個名為 `getRandomComputerChoice` 的函式（程式），它以可用選項的陣列作為參數，並傳回對所選問題的隨機答案。

```js
assert.isFunction(getRandomComputerChoice);
const randomArray = ['a', 'b', 'c'];
assert.oneOf(getRandomComputerChoice(randomArray), randomArray)
```

你應該有一個名為 `getResults` 的函式（程式）。

```js
assert.isFunction(getResults);
```

你的 `getResults` 函式（程式）應該將問題物件作為第一個參數，將電腦的選擇作為第二個參數。

```js
const testQuestion = questions[0];
const wrongChoice = questions[0].choices.find(choice => choice !== questions[0].answer);
assert.equal(getResults(testQuestion, testQuestion.answer), `The computer's choice is correct!`);
assert.equal(getResults(testQuestion, wrongChoice), `The computer's choice is wrong. The correct answer is: ${testQuestion.answer}`);
```

如果電腦選擇與答案相符，`getResults` 應該傳回 `The computer's choice is correct!`

```js
assert.equal(getResults(questions[0], questions[0].answer), `The computer's choice is correct!`);
```

如果電腦的選擇與答案不符，`getResults` 應該傳回 `The computer's choice is wrong. The correct answer is: <correct-answer>`，其中 `<correct-answer>` 是所選問題的正確答案值。

```js
assert.equal(getResults({category: 'misc', choices: ['a', 'b', 'c'], question: "question?", answer: "b"}, "a"), `The computer's choice is wrong. The correct answer is: b`)
```

你的 `getResults` 函式應該使用精確的相等性比較，而不是子字串匹配。

```js
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Hamburger"}, "Ham"), `The computer's choice is wrong. The correct answer is: Hamburger`);
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Ham"}, "Hamburger"), `The computer's choice is wrong. The correct answer is: Ham`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const questions = [
    {
        category: "science",
        question: "What is the chemical symbol for potassium?",
        choices: ["P", "K", "Pt"],
        answer: "K"
    },
    {
        category: "science",
        question: "What is the unit of electrical resistance?",
        choices: ["Ohm", "Coulomb", "Sievert"],
        answer: "Ohm"
    },
    {
        category: "geography",
        question: "What is the capital city of Australia?",
        choices: ["Sidney", "Canberra", "Wellington"],
        answer: "Canberra"
    },
    {
        category: "literature",
        question: 'Who wrote "1984"?',
        choices: ["Ray Bradbury", "Aldous Huxley", "George Orwell"],
        answer: "George Orwell"
    },
    {
        category: "sport",
        question: "How many players are on a standard volleyball team?",
        choices: ["6", "7", "12"],
        answer: "6"
    }
]

function getRandomQuestion(questionList) {
    const randomIndex = Math.floor(Math.random() * questionList.length);
    return questionList[randomIndex];
}

function getRandomComputerChoice(choices) {
    const randomIndex = Math.floor(Math.random() * choices.length);
    return choices[randomIndex];
}

function getResults(obj, computerChoice) {
    return computerChoice === obj.answer
        ? "The computer's choice is correct!"
        : `The computer's choice is wrong. The correct answer is: ${obj.answer}`;
}

const questionObj = getRandomQuestion(questions);
const { question, choices } = questionObj;

console.log(question);
console.log(`Choices: ${choices}`);

const computerChoice = getRandomComputerChoice(choices);
console.log(`Computer chooses: ${computerChoice}`);

const results = getResults(questionObj, computerChoice);
console.log(results);

```
