---
id: 657a0ea50da0c8d9d6d7950a
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このプロジェクトでは、サイコロゲームを作りながらアルゴリズム思考を学びます。全部で6ラウンドあり、各ラウンドでプレイヤーは最大3回サイコロを振ってスコアを獲得できます。

HTMLとCSSは用意されています。自由に確認してください。準備ができたら、HTMLの変数を設定しましょう。すべての`.die`要素を取得して`listOfAllDice`変数に代入してください。スコアの入力（`input`要素が`#score-options`内にあります）とスコアのスパンを取得し、それぞれ`scoreInputs`と`scoreSpans`に代入してください。`#current-round`要素を`roundElement`に、`#current-round-rolls`要素を`rollsElement`に代入し、同様に`#total-score`と`#score-history`要素も代入してください。`#roll-dice-btn`、`#keep-score-btn`、`#rules-btn`、`.rules-container`も適切な名前の変数に代入してください。

ユーザーが`Show rules`ボタンをクリックすると、ゲームルールの表示と非表示を切り替えられるようにしてください。ゲームルールの状態を追跡するために`isModalShowing`変数を作成してください。

ユーザーがサイコロを振るたびに、すべてのサイコロの値を追跡する必要があります。これを追跡するために`diceValuesArr`変数を作成してください。

ゲームの進行中は、現在のスコア、合計スコア、振った回数、プレイヤーがどのラウンドにいるかを追跡する必要があります。これらの目的のために`rolls`、`score`、`round`変数を宣言してください。

これらの変数の初期値が何になるべきか考えてみましょう。すべての変数は再代入可能である必要があります。

# --hints--

`.die`要素を`listOfAllDice`に代入してください。

```js
assert.deepEqual(listOfAllDice, document.querySelectorAll(".die"));
```

`#score-options input`要素を`scoreInputs`に代入してください。

```js
assert.deepEqual(scoreInputs, document.querySelectorAll("#score-options input"));
```

`#score-options span`要素を`scoreSpans`に代入してください。

```js
assert.deepEqual(scoreSpans, document.querySelectorAll("#score-options span"));
```

`#current-round`要素を`roundElement`に代入してください。

```js
assert.deepInclude([document.getElementById("current-round"), document.querySelector("#current-round")], roundElement);
```

`#current-round-rolls`要素を`rollsElement`に代入してください。

```js
assert.deepInclude([document.getElementById("current-round-rolls"), document.querySelector("#current-round-rolls")], rollsElement);
```

`#total-score`要素を`totalScoreElement`に代入してください。

```js
assert.deepInclude([document.getElementById("total-score"), document.querySelector("#total-score")], totalScoreElement);
```

`#score-history`要素を`scoreHistory`に代入してください。

```js
assert.deepInclude([document.getElementById("score-history"), document.querySelector("#score-history")], scoreHistory);
```

`#roll-dice-btn`要素を`rollDiceBtn`に代入してください。

```js
assert.deepInclude([document.getElementById("roll-dice-btn"), document.querySelector("#roll-dice-btn")], rollDiceBtn);
```

`#keep-score-btn`要素を`keepScoreBtn`に代入してください。

```js
assert.deepInclude([document.getElementById("keep-score-btn"), document.querySelector("#keep-score-btn")], keepScoreBtn);
```

`#rules-btn`要素を`rulesBtn`に代入してください。

```js
assert.deepInclude([document.getElementById("rules-btn"), document.querySelector("#rules-btn")], rulesBtn);
```

`.rules-container`要素を`rulesContainer`に代入してください。

```js
assert.deepEqual(rulesContainer, document.querySelector(".rules-container"));
```

`isModalShowing`変数は`false`の値を持つべきです。

```js
assert.isFalse(isModalShowing);
```

`diceValuesArr`は空の配列を値として持つべきです。

```js
assert.isArray(diceValuesArr);
assert.isEmpty(diceValuesArr);
```

`rolls`変数は`0`の値を持つべきです。

```js
assert.strictEqual(rolls, 0);
```

`score`変数は`0`の値を持つべきです。

```js
assert.strictEqual(score, 0);
```

`round`変数は`1`の値を持つべきです。

```js
assert.strictEqual(round, 1);
```

すべての新しい変数は`let`で宣言してください。

```js
assert.match(code, /let isModalShowing/);
assert.match(code, /let diceValuesArr/);
assert.match(code, /let rolls/);
assert.match(code, /let score/);
assert.match(code, /let round/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Advanced Dice Game</title>
  <link rel="stylesheet" href="styles.css" />
</head>

<body>
  <header>
    <h1>Advanced Dice Game</h1>
    <button class="btn" id="rules-btn" type="button">Show rules</button>
    <div class="rules-container">
      <h2>Rules</h2>
      <ul>
        <li>There are total of six rounds</li>
        <li>You can only roll the dice three times per round</li>
        <li>To start the game, roll the dice</li>
        <li>
          Then, choose from one of the selected scores or roll the dice again
        </li>
        <li>
          If you choose a selected score, then you will move to the next round
        </li>
        <li>
          If you decline to choose a selected score, then you can roll the
          dice again two more times
        </li>
      </ul>
      <h2 class="points">Points</h2>
      <ul>
        <li>Three of a kind: Sum of all five dice</li>
        <li>Four of a kind: Sum of all five dice</li>
        <li>Full house: Three of a kind and a pair - 25 points</li>
        <li>
          Small straight: Four of the dice have consecutive values - 30 points
        </li>
        <li>
          Large straight: All five dice have consecutive values - 40 points
        </li>
      </ul>
    </div>
  </header>

  <main>
    <form id="game">
      <div id="dice">
        <div class="die"></div>
        <div class="die"></div>
        <div class="die"></div>
        <div class="die"></div>
        <div class="die"></div>
      </div>

      <p class="rounds-text">
        <strong>Rolls:</strong> <span id="current-round-rolls">0</span> |
        <strong>Round:</strong> <span id="current-round">1</span>
      </p>

      <div id="score-options">
        <div>
          <input type="radio" name="score-options" id="three-of-a-kind" value="three-of-a-kind" disabled />
          <label for="three-of-a-kind">Three of a kind<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="four-of-a-kind" value="four-of-a-kind" disabled />
          <label for="four-of-a-kind">Four of a kind<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="full-house" value="full-house" disabled />
          <label for="full-house">Full house<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="small-straight" value="small-straight" disabled />
          <label for="small-straight">Small straight<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="large-straight" value="large-straight" disabled />
          <label for="large-straight">Large straight<span></span></label>
        </div>

        <div>
          <input type="radio" name="score-options" id="none" value="none" disabled />
          <label for="none">None of the above<span></span></label>
        </div>
      </div>

      <button class="btn" id="keep-score-btn" type="button">
        Keep the above selected score
      </button>
      <button class="btn" id="roll-dice-btn" type="button">
        Roll the dice
      </button>
    </form>

    <div id="scores">
      <h3>Score history (Total score: <span id="total-score">0</span>)</h3>
      <ol id="score-history"></ol>
    </div>
  </main>
  <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --black: #000;
  --white: #fff;
  --grey: #3b3b4f;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

header {
  color: var(--light-grey);
  text-align: center;
}

h1 {
  font-size: 2.5rem;
  margin: 25px 0;
}

.rules-container {
  display: none;
  background-color: var(--light-grey);
  color: var(--black);
  width: 50%;
  margin: 20px auto;
  height: 300px;
  border-radius: 10px;
  overflow-y: scroll;
}

.rules-container ul {
  list-style-type: none;
}

.points {
  margin-top: 15px;
}

main {
  background-color: var(--light-grey);
  padding: 20px;
  display: grid;
  grid-template-columns: repeat(auto-fit, minmax(500px, 1fr));
  gap: 1rem;
  margin: auto;
  justify-items: center;
  width: 50%;
  border-radius: 10px;
}

#dice {
  display: flex;
  justify-content: space-around;
  margin-bottom: 15px;
}

.die {
  width: 40px;
  height: 40px;
  text-align: center;
  margin-right: 15px;
  border: 4px solid var(--black);
  padding: 10px;
}

.rounds-text {
  text-align: center;
}

input[type="radio"]:disabled + label {
  color: var(--grey);
}

#score-history {
  margin-top: 15px;
  text-align: center;
  list-style-position: inside;
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px 0 10px 0.5rem;
  color: var(--black);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover:enabled {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

@media (max-width: 992px) {
  main {
    width: 100%;
  }
}

@media (max-width: 500px) {
  .btn {
    width: 120px;
  }
}

```

```js
--fcc-editable-region--

--fcc-editable-region--
```
