---
id: a6b0bb188d873cb2c8729495
title: Einen HTML-Entity-Konverter implementieren
challengeType: 26
dashedName: implement-an-html-entity-converter
---

# --description--

In diesem Labor geht es darum, Sonderzeichen in einem String durch ihre entsprechenden HTML-Entities zu ersetzen.

**Ziel**: Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine `convertHTML`-Funktion haben, die einen String als Argument annimmt.
1. Die `convertHTML`-Funktion sollte einen neuen String zurückgeben, indem sie Sonderzeichen im Argument-String in die entsprechenden HTML-Entities umwandelt.

   - `&` sollte in `&amp;` umgewandelt werden.
   - `<` sollte in `&lt;` umgewandelt werden.
   - `>` sollte in `&gt;` umgewandelt werden.
   - `"` sollte in `&quot;` umgewandelt werden.
   - `'` sollte in `&apos;` umgewandelt werden.

# --hints--

Sie sollten eine `convertHTML`-Funktion haben.

```js
assert.isFunction(convertHTML);
```

`convertHTML("Dolce & Gabbana")` sollte den String `Dolce &amp; Gabbana` zurückgeben.

```js
assert.match(convertHTML('Dolce & Gabbana'), /Dolce &amp; Gabbana/);
```

`convertHTML("Hamburgers < Pizza < Tacos")` sollte den String `Hamburgers &lt; Pizza &lt; Tacos` zurückgeben.

```js
assert.match(
  convertHTML('Hamburgers < Pizza < Tacos'),
  /Hamburgers &lt; Pizza &lt; Tacos/
);
```

`convertHTML("Sixty > twelve")` sollte den String `Sixty &gt; twelve` zurückgeben.

```js
assert.match(convertHTML('Sixty > twelve'), /Sixty &gt; twelve/);
```

`convertHTML('Stuff in "quotation marks"')` sollte den String `Stuff in &quot;quotation marks&quot;` zurückgeben.

```js
assert.match(
  convertHTML('Stuff in "quotation marks"'),
  /Stuff in &quot;quotation marks&quot;/
);
```

`convertHTML("Schindler's List")` sollte den String `Schindler&apos;s List` zurückgeben.

```js
assert.match(convertHTML("Schindler's List"), /Schindler&apos;s List/);
```

`convertHTML("<>")` sollte den String `&lt;&gt;` zurückgeben.

```js
assert.match(convertHTML('<>'), /&lt;&gt;/);
```

`convertHTML("abc")` sollte den String `abc` zurückgeben.

```js
assert.strictEqual(convertHTML('abc'), 'abc');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const map = {
  '&': '&amp;',
  '<': '&lt;',
  '>': '&gt;',
  '"': '&quot;',
  "'": '&apos;'
};

function convertHTML(str) {
  return str.replace(/[&<>"']/g, function(char) {
    return map[char];
  });
}
```
