---
id: 66ea7fc048a6e26e33c598fe
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

حان الوقت الآن لتوليد رقم عشوائي باستخدام الطريقة `Math.random()`.

أنشئ متغيرًا اسمه `randomNum` وقم بتعيين قيمة نتيجة استدعاء الطريقة `Math.random()` له.

بعد ذلك، سجّل المتغير `randomNum` في وحدة التحكم.

جرّب إضافة مسافة في الشفرة لإعادة تشغيل الروبوت ورؤية أرقام عشوائية مختلفة يتم توليدها.

# --hints--

يجب أن يكون لديك متغير اسمه `randomNum`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.randomNum);
```

يجب تعيين نتيجة استدعاء الطريقة `Math.random()` إلى المتغير `randomNum`.

```js
const explorer = await __helpers.Explorer(code);
const randomNum = explorer.variables.randomNum;
assert.equal(randomNum?.value?.toString(), "Math.random()");
```

يجب أن تسجل المتغير `randomNum` في وحدة التحكم.

```js
assert.match(code, /console\.log\(\s*randomNum\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");
--fcc-editable-region--

--fcc-editable-region--
```
