---
id: 66ea8239b9eeb76f12299030
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

الجزء التالي من ورشة العمل هو مراجعة كيفية توليد رقم عشوائي بين قيمتين.

ابدأ بإضافة `console.log()` آخر يسجل الرسالة `"Now, generate a random number between two values."` في وحدة التحكم.

ثم، أنشئ متغيرًا اسمه `min` وقم بتعيين القيمة `1` له، وأنشئ متغيرًا اسمه `max` وقم بتعيين القيمة `100` له.

في الخطوة التالية، ستولد رقمًا عشوائيًا بين هاتين القيمتين.

# --hints--

يجب أن يكون لديك `console.log()` يسجل الرسالة `"Now, generate a random number between two values."`.

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+number\s+between\s+two\s+values\.\s*\1\);?/);
```

يجب أن يكون لديك متغير اسمه `min`.

```js
assert.isNotNull(min);
```

يجب أن يكون متغير `min` رقمًا.

```js
assert.isNumber(min);
```

يجب تعيين الرقم `1` إلى متغير `min`.

```js
assert.equal(min, 1);
```

يجب أن يكون لديك متغير اسمه `max`.

```js
assert.isNotNull(max);
```

يجب أن يكون متغير `max` رقمًا.

```js
assert.isNumber(max);
```

يجب تعيين الرقم `100` إلى متغير `max`.

```js
assert.equal(max, 100); 
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);
--fcc-editable-region--

--fcc-editable-region--
```
