---
id: 66ea7adae8053065a64f9002
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستراجع العمل مع الطرق المختلفة لكائن `Math` من خلال بناء Mathbot. سيُسجل هذا Mathbot بعض العمليات الرياضية والرسائل في وحدة التحكم بهدف تعليمك عن كائن `Math` في JavaScript.

في هذه الخطوة الأولى، ابدأ بإنشاء متغير يسمى `botName` وقم بتعيين القيمة النصية `"MathBot"` له.

ثم، أنشئ متغيرًا آخر يسمى `greeting` وقم بتعيين الجملة `"Hi there! My name is [botName goes here] and I am here to teach you about the Math object!"`. بدلاً من عنصر نائب `[botName goes here]`، استخدم المتغير `botName`.

يمكنك استخدام قوالب حرفية أو ربط السلاسل النصية باستخدام عامل `+` لتحقيق ذلك.

أخيرًا، سجّل المتغير `greeting` في وحدة التحكم.

# --hints--

يجب أن يكون لديك متغير اسمه `botName`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.botName);
```

يجب أن يكون المتغير `botName` سلسلة نصية.

```js
assert.isString(botName);
```

يجب تعيين القيمة النصية `"MathBot"` للمتغير `botName`.

```js
assert.strictEqual(botName, "MathBot");
```

يجب أن يكون لديك متغير يسمى `greeting`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.greeting);
```

يجب أن يكون المتغير `greeting` سلسلة نصية.

```js
assert.isString(greeting);
```

يجب تعيين الجملة التالية للمتغير `greeting`: `"Hi there! My name is [MathBot goes here] and I am here to teach you about the Math object!"`. تأكد من استبدال عنصر نائب `[MathBot goes here]` بالمتغير `botName`.

```js
assert.strictEqual(greeting, "Hi there! My name is MathBot and I am here to teach you about the Math object!");
```

يجب أن تسجل المتغير `greeting` في وحدة التحكم.

```js
assert.match(code, /console\.log\(greeting\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
