---
id: 66ea845f54195c7020cc521b
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

الصيغة لتوليد رقم عشوائي بين قيمتين هي كما يلي:

```js
Math.random() * (maximum - minimum) + minimum;
```

هذا سينتج قيمة عدد عشري بين قيمتين.

أنشئ متغيرًا اسمه `randomNum2` وقم بتعيين نتيجة توليد قيمة بين القيمتين `min` و `max` له.

ثم سجّل المتغير `randomNum2` لرؤية النتيجة. جرّب تحديث الصفحة لرؤية نتائج مختلفة.

# --hints--

يجب أن يكون لديك متغير اسمه `randomNum2`.

```js
assert.isNotNull(randomNum2);
```

يجب أن يكون المتغير `randomNum2` عددًا.

```js
assert.isNumber(randomNum2);
```

يجب تعيين نتيجة توليد قيمة بين `min` و `max`. تأكد من عدم تعيين قيمة ثابتة. استخدم الصيغة المقدمة أعلاه.

```js
assert.isAtLeast(randomNum2, min);
assert.isAtMost(randomNum2, max);

assert.notMatch(code, /randomNum2\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum2\s*=\s*\d+;?/);
```

يجب أن تسجل المتغير `randomNum2` في وحدة التحكم.

```js
assert.match(code, /console\.log\(\s*randomNum2\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;
--fcc-editable-region--

--fcc-editable-region--
```
