---
id: 66eb488c0129e192542f0865
title: الخطوة 11
challengeType: 1
dashedName: step-11
---

# --description--

في الدروس السابقة، تعلمت كيف تتعامل مع الطريقة `Math.ceil()` بهذا الشكل:

```js
const price = 10.01;
Math.ceil(price); // 11
```

الطريقة `Math.ceil()` تُقرب القيمة إلى أقرب عدد صحيح أكبر أو يساويها.

أنشئ متغيرًا اسمه `numRoundedUp` وقم بتعيين نتيجة تقريب الرقم العشري `3.2` إلى أقرب عدد صحيح أكبر أو يساويه.

ثم، سجّل المتغير `numRoundedUp` في وحدة التحكم لرؤية النتيجة.

# --hints--

يجب أن يكون لديك متغير اسمه `numRoundedUp`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.numRoundedUp);
```

يجب أن يكون المتغير `numRoundedUp` رقمًا.

```js
assert.isNumber(numRoundedUp);
```

يجب أن يساوي متغيرك `numRoundedUp` نتيجة تقريب الرقم العشري `3.2` إلى أقرب عدد صحيح أكبر أو يساويه.

```js
assert.equal(numRoundedUp, 4);
```

لا يجب أن تحدد القيمة `4` مباشرةً للمتغير `numRoundedUp`.

```js
const explorer = await __helpers.Explorer(code);
const numRoundedUp = explorer.variables.numRoundedUp;
assert.notEqual(numRoundedUp?.value?.toString(), "4");
```

يجب أن تسجل المتغير `numRoundedUp` في وحدة التحكم.

```js
assert.match(code, /console\.log\(\s*numRoundedUp\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

--fcc-editable-region--

--fcc-editable-region--
```
