---
id: 66eb3ec3b07dfe8f898646c0
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

كما تعلمت في الدروس السابقة، تقوم الطريقة `Math.floor()` بتقريب القيمة إلى أقرب عدد صحيح أقل.

```js
const price = 10.99;
Math.floor(price); // 10
```

أنشئ متغيرًا اسمه `numRoundedDown` وقم بتعيين نتيجة تقريب الرقم العشري `6.7` إلى أقرب عدد صحيح أقل.

ثم، سجّل المتغير `numRoundedDown` في وحدة التحكم لرؤية النتيجة.

# --hints--

يجب أن يكون لديك متغير اسمه `numRoundedDown`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.numRoundedDown);
```

يجب أن يكون المتغير `numRoundedDown` رقمًا.

```js
assert.isNumber(numRoundedDown);
```

يجب أن يساوي متغير `numRoundedDown` نتيجة تقريب الرقم العشري `6.7` إلى أقرب عدد صحيح أقل.

```js
assert.equal(numRoundedDown, 6);
```

يجب ألا تحدد القيمة `6` مباشرةً في المتغير `numRoundedDown`.

```js
const explorer = await __helpers.Explorer(code);
const numRoundedDown = explorer.variables.numRoundedDown;
assert.notEqual(numRoundedDown?.value?.toString(), "6");
```

يجب أن تسجل المتغير `numRoundedDown` في وحدة التحكم.

```js
assert.match(code, /console\.log\(\s*numRoundedDown\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");
--fcc-editable-region--

--fcc-editable-region--
```
