---
id: 66ece147323c91f654bf8dd9
title: الخطوة 15
challengeType: 1
dashedName: step-15
---

# --description--

في الدروس السابقة، تعلمت كيف تحصل على القيم القصوى والحد الأدنى من مجموعة أرقام مثل هذه:

```js
Math.max(1, 2, 3, 4, 5); // 5
Math.min(1, 2, 3, 4, 5); // 1
```

أنشئ متغيرًا يسمى `maxNum` وقم بتعيين نتيجة إيجاد الرقم الأكبر بين الأرقام `3, 125, 55, 24`. ثم، سجّل قيمة `maxNum` في وحدة التحكم.

أسفل ذلك، أنشئ متغيرًا يسمى `minNum` وقم بتعيين نتيجة إيجاد الرقم الأصغر بين الأرقام `6, 90, 14, 90, 2`. ثم، سجّل قيمة `minNum` في وحدة التحكم.

# --hints--

يجب أن يكون لديك متغير يسمى `maxNum`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.maxNum);
```

يجب تعيين متغير `maxNum` نتيجة إيجاد الرقم الأكبر بين الأرقام `3, 125, 55, 24`.

```js
assert.equal(maxNum, 125);
```

لا يجب أن تُدخل القيمة `125` مباشرة في متغير `maxNum`. تأكد من استخدام طريقة `Math.max()` بدلاً من ذلك.

```js
const explorer = await __helpers.Explorer(code);
assert.notEqual(explorer.variables.maxNum?.value?.toString(), "125");
```

يجب أن تسجل قيمة `maxNum` في وحدة التحكم.

```js
assert.match(code, /console\.log\(\s*maxNum\s*\)/);
```

يجب أن يكون لديك متغير يسمى `minNum`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.minNum);
```

يجب تعيين متغير `minNum` نتيجة إيجاد الرقم الأصغر بين الأرقام `6, 90, 14, 90, 2`.

```js
assert.equal(minNum, 2);
```

لا يجب أن تُدخل القيمة `2` مباشرة في متغير `minNum`. تأكد من استخدام طريقة `Math.min()` بدلاً من ذلك.

```js
const explorer = await __helpers.Explorer(code);
assert.notEqual(explorer.variables.minNum?.value?.toString(), "2");
```

يجب أن تسجل قيمة `minNum` في وحدة التحكم.

```js
assert.match(code, /console\.log\(\s*minNum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

console.log(
  "The Math.round() method rounds the value to the nearest whole integer."
);

const numRounded = Math.round(2.7);
console.log(numRounded);
const numRounded2 = Math.round(11.2);
console.log(numRounded2);

console.log("The Math.max() and Math.min() methods are used to get the maximum and minimum number from a range.");

--fcc-editable-region--

--fcc-editable-region--
```
