---
id: 66ece0b518f5dcf5930612ab
title: الخطوة 14
challengeType: 1
dashedName: step-14
---

# --description--

في الجزء الأخير من ورشة العمل، ستراجع كيفية الحصول على القيم الحد الأدنى والحد الأقصى لمجموعة من الأرقام.

ابدأ بإضافة `console.log()` يسجل الرسالة `"The Math.max() and Math.min() methods are used to get the maximum and minimum number from a range."` في وحدة التحكم.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن تسجل الرسالة `"The Math.max() and Math.min() methods are used to get the maximum and minimum number from a range."` في وحدة التحكم.

```js
assert.equal(
  getLogs()[14],
  "The Math.max() and Math.min() methods are used to get the maximum and minimum number from a range."
);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

console.log(
  "The Math.round() method rounds the value to the nearest whole integer."
);

const numRounded = Math.round(2.7);
console.log(numRounded);
const numRounded2 = Math.round(11.2);
console.log(numRounded2);
--fcc-editable-region--

--fcc-editable-region--
```
