---
id: 69218e46f4490b33f2481a1a
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

الصيغة لتوليد عدد صحيح عشوائي بين قيمتين هي كما يلي:

```js
Math.floor(Math.random() * (maximum - minimum) + minimum);
```

هذا سينتج قيمة عدد صحيح بين قيمتين.

أنشئ متغيرًا اسمه `randomInt` وقم بتعيين نتيجة توليد قيمة بين القيمتين `min` و `max` له.

ثم سجّل المتغير `randomInt` لرؤية النتيجة. جرّب تحديث الصفحة لرؤية نتائج مختلفة.

# --hints--

يجب أن يكون لديك متغير اسمه `randomInt`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.randomInt);
```

يجب أن يكون المتغير `randomInt` رقمًا.

```js
assert.isNumber(randomInt);
```

يجب تعيين نتيجة توليد قيمة بين القيمتين `min` و `max`. تأكد من عدم تعيين قيمة ثابتة. يجب أن تستخدم الصيغة المقدمة أعلاه.

```js
assert.isAtLeast(randomInt, min);
assert.isAtMost(randomInt, max);

const explorer = await __helpers.Explorer(code);  
const expression = explorer.variables.randomInt?.value.toString();  
assert.include(expression, "Math.floor");  
assert.include(expression, "Math.random()");  
assert.match(expression, /\bmin\b/);  
assert.match(expression, /\bmax\b/);  
```

يجب أن تسجل المتغير `randomInt` في وحدة التحكم.

```js
assert.match(code, /console\.log\(\s*randomInt\s*\);?/);
```


# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");
--fcc-editable-region--

--fcc-editable-region--
```
