---
id: 69860157f7fce39f9df399a3
title: N-퀸 알고리즘 구현하기
challengeType: 26
dashedName: implement-the-n-queens-algorithm
---

# --description--

N-퀸 문제는 N×N 체스판 위에 N개의 퀸을 서로 공격하지 못하도록(같은 행, 열, 대각선에 두지 않도록) 배치하는 문제입니다.

예를 들어, 4x4 체스판이 있을 때, 유효한 배치 중 하나는 다음과 같습니다:

```md
[1, 3, 0, 2]
```

이는 0행에 퀸이 1열에, 1행에 퀸이 3열에, 2행에 퀸이 0열에, 3행에 퀸이 2열에 배치되었음을 의미합니다.

시각적으로 이 배치는 다음과 같습니다:

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

여기서 `Q`는 퀸을, `.`는 빈 칸을 나타냅니다.

이 실습에서는 깊이 우선 탐색 방식을 사용해 N-퀸 문제 해결기를 구현할 것입니다.


**목표**: 아래 사용자 스토리를 모두 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `dfsNQueens`라는 이름의 함수가 있어야 합니다.
2. 함수는 정확히 하나의 인수, 정수 `n`를 받아야 합니다.
3. `n`가 `1`보다 작으면 함수는 빈 배열(`[]`)을 반환해야 합니다.
4. 함수는 해결책 목록을 반환해야 하며, 각 해결책은 길이 `n`인 배열이고, 인덱스 `i`의 요소는 `i` 행에 있는 퀸의 0부터 시작하는 열 인덱스입니다.

# --hints--

하나의 인수를 받는 `dfsNQueens`라는 함수가 있어야 합니다.

```js
assert.isFunction(dfsNQueens);
assert.lengthOf(dfsNQueens, 1);
```

`n`가 `1`보다 작으면 함수는 빈 배열을 반환해야 합니다.

```js
assert.sameDeepOrderedMembers(dfsNQueens(0), []);
assert.sameDeepOrderedMembers(dfsNQueens(-1), []);
assert.sameDeepOrderedMembers(dfsNQueens(-5), []);
```

함수는 각 해결책이 길이 `n`인 배열인 해결책 배열을 반환해야 합니다.

```js
const output = dfsNQueens(4);
assert.isArray(output);
output.forEach(solution => {
    assert.isArray(solution);
    assert.lengthOf(solution, 4);
})
```

`dfsNQueens(1)`는 `[[0]]`를 반환해야 합니다.

```js
assert.sameDeepOrderedMembers(dfsNQueens(1), [[0]]);
```

`dfsNQueens(2)`는 `[]`를 반환해야 합니다.

```js
assert.sameDeepOrderedMembers(dfsNQueens(2), []);
```

`dfsNQueens(3)`는 `[]`를 반환해야 합니다.

```js
assert.sameDeepOrderedMembers(dfsNQueens(3), []);
```

`dfsNQueens(4)`는 `[[1, 3, 0, 2], [2, 0, 3, 1]]`를 반환해야 합니다.

```js
assert.sameDeepOrderedMembers(dfsNQueens(4), [[1, 3, 0, 2], [2, 0, 3, 1]]);
```

`dfsNQueens(5)`는 `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`를 반환해야 합니다.

```js
assert.sameDeepOrderedMembers(dfsNQueens(5), [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]);
```

`dfsNQueens(5).length`는 `10`여야 합니다.

```js
assert.lengthOf(dfsNQueens(5), 10);
```

`dfsNQueens(8).length`는 `92`여야 합니다.

```js
assert.lengthOf(dfsNQueens(8), 92);
```

`dfsNQueens`는 올바른 결과를 반환해야 합니다.

```js
assert.lengthOf(dfsNQueens(7), 40);
assert.sameDeepOrderedMembers(dfsNQueens(6), [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfsNQueens(n) {
  if (n < 1) return [];

  const results = [];

  function isSafe(queens, row, col) {
    for (let r = 0; r < queens.length; r++) {
      const c = queens[r];
      if (c === col || Math.abs(row - r) === Math.abs(col - c)) {
        return false;
      }
    }
    return true;
  }

  function dfs(queens) {
    const row = queens.length;
    if (row === n) {
      results.push(queens.slice());
      return;
    }

    for (let col = 0; col < n; col++) {
      if (isSafe(queens, row, col)) {
        queens.push(col);
        dfs(queens);
        queens.pop();
      }
    }
  }

  dfs([]);
  return results;
}
```
