---
id: 615f2abbe7d18d49a1e0e1c8
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

基本的なHTMLのボイラープレートを用意しました。

`h1`要素の中に`body`要素を作成し、テキストに`Nutrition Facts`を入れてください。

# --hints--

新しい`h1`要素を追加してください。

```js
assert.exists(document.querySelector('h1'));
```

`h1`要素は`body`要素の中にあるべきです。

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

`h1`要素にはテキスト`Nutrition Facts`を入れてください。

```js
assert(document.querySelector('h1')?.innerText === 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>

</body>
</html>
--fcc-editable-region--
```

```css

```
