---
id: 615f34948891834dd77655a6
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

テキストが `p` の2つ目の `Serving size 2/3 cup (55g)` 要素を追加してください。

# --hints--

2つ目の `p` 要素があるはずです。

```js
assert(document.querySelectorAll('p')?.length === 2);
```

2つ目の `p` 要素は `body` 要素の中にあるはずです。

```js
assert(document.querySelectorAll('p')?.[1]?.parentElement?.localName === 'body');
```

2つ目の `p` 要素は既存の `p` 要素の後に来るはずです。

```js
assert(document.querySelectorAll('p')?.[1]?.previousElementSibling?.localName === 'p');
```

2つ目の `p` 要素のテキストは `Serving size 2/3 cup (55g)` であるべきです。

```js
assert(document.querySelectorAll('p')?.[1]?.innerText === 'Serving size 2/3 cup (55g)');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
    
</body>
</html>
--fcc-editable-region--
```

```css

```
