---
id: 615f3cafd794015aa9547a6d
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

`h1`、`h2` などのタグの使用はHTMLの意味構造を決定することを覚えておいてください。ただし、これらの要素のCSSを調整して、視覚的な流れや階層をコントロールできます。

`h1` のルールを作成し、`font-weight` プロパティを `800` に設定してください。これにより、`h1` のテキストが太字になります。

# --hints--

`h1` セレクターを作成してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

`h1` セレクターには `font-weight` プロパティを `800` に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === '800');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--

--fcc-editable-region--
```
