---
id: 615f2d4150fe0d4cbd0f2628
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`h1` 要素の下に、テキスト `p` を持つ `8 servings per container` 要素を追加してください。

# --hints--

新しい `p` 要素を追加してください。

```js
assert.exists(document.querySelector('p'));
```

`p` 要素は `body` 要素の中にあるべきです。

```js
assert(document.querySelector('p')?.parentElement?.localName === 'body');
```

`p` 要素は `h1` 要素の後に来るべきです。

```js
assert(document.querySelector('p')?.previousElementSibling?.localName === 'h1');
```

`p` 要素はテキスト `8 servings per container` を持つべきです。

```js
assert(document.querySelector('p')?.innerText === '8 servings per container');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
    
</body>
</html>
--fcc-editable-region--
```

```css

```
