---
id: 685523c05fb41d0848c4b52f
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro do primeiro elemento `div`, crie um elemento `span` com o texto `Soundflow`, um elemento `h1` com o texto `Discover New Music` e um elemento `p` com o texto `Stream your favorite tracks and discover new artists.`.

# --hints--

Você deve criar um elemento `span` com o texto `Soundflow` dentro do primeiro elemento `div`.

```js
assert.equal(document.querySelector("div > div:first-child > span")?.textContent, "Soundflow")
```

Você deve criar um elemento `h1` com o texto `Discover New Music` dentro do primeiro elemento `div`.

```js
assert.equal(document.querySelector("div > div:first-child > h1")?.textContent, "Discover New Music")
```

Você deve criar um elemento `p` com o texto `Stream your favorite tracks and discover new artists.` dentro do primeiro elemento `div`.

```js
assert.equal(document.querySelector("div > div:first-child > p")?.textContent, "Stream your favorite tracks and discover new artists.")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </div>
      <div></div>
    </div>
  </body>
</html>
```
