---
id: 68551b677a5f6898af364b85
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você aprenderá mais sobre Tailwind CSS usando suas classes utilitárias de flexbox para construir um componente de call to action (CTA) responsivo.

O boilerplate básico de HTML foi fornecido para você, então você precisa construir os elementos HTML primeiro.

Comece criando um elemento `div`. Dentro desse `div`, crie mais dois elementos `div` separados.

# --hints--

Você deve criar um elemento `div`.

```js
assert.exists(document.querySelector("div"));
```

Você deve ter mais dois elementos `div` dentro do elemento `div`.

```js
assert.lengthOf(document.querySelectorAll("div > div"), 2)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
