---
id: 68552464ef378e1051a55c28
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Dentro do segundo `div`, crie dois elementos anchor.O primeiro deve ter o texto `Learn more` e o segundo deve ter o texto `Start listening`.

Para o atributo `href` dos dois elementos anchor, defina ambos como `#`.

# --hints--

Você deve criar dois elementos `a` dentro do segundo elemento `div`

```js
assert.lengthOf(document.querySelectorAll("div > div:nth-child(2) > a"), 2)
```

Seu primeiro elemento `a` deve ter o texto `Learn more` e um atributo `href` com o valor `#` dentro do segundo elemento `div`.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[0]?.textContent, "Learn more")
assert.equal(links[0]?.getAttribute("href"), "#")
```

Seu segundo elemento `a` deve ter o texto `Start listening` e um atributo `href` com valor `#` dentro do segundo elemento `div`.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[1]?.textContent, "Start listening")
assert.equal(links[1]?.getAttribute("href"), "#")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
