---
id: 6858f3a67aa8580ff37016c7
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Se você ajustar a largura da aba de visualização para a largura idêntica a de um dispositivo móvel novamente, você notará que o conteúdo está sempre lado a lado.Isso não é o que você quer em dispositivos móveis.

Para corrigir isso, adicione a classe `flex-col` ao `div`.Isso vai mudar a direção do flex para coluna, então os elementos vão se empilhar uns sobre os outros.

Aqui está como usar a classe utilitária `flex-col`:

```html
<div class="flex flex-col">
  <element></element>
  <element></element>
</div>
```

# --hints--

Seu elemento `div` de abertura deve ter a classe utilitária `flex-col`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex-col"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
