---
id: 6858ef974c8ad5dcd7abb9f1
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

O conteúdo dentro do `div` pode ser melhor espaçado e organizado.Para o restante do espaçamento e desse arranjo, você precisa usar flexbox.

Adicione a classe `flex` ao `div` para que os dois elementos filhos `div` fiquem lado a lado.

Além disso, adicione `justify-around` para que o espaço ao redor e entre o conteúdo seja distribuído uniformemente e `items-center` para que todo o conteúdo fique finalmente centralizado no `div`. 

# --hints--

Seu elemento externo `div` deve ter a classe utilitária `flex`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex"));
```

Seu elemento externo `div` deve ter a classe utilitária `justify-around`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("justify-around"));
```

Seu elemento externo `div` deve ter a classe utilitária `items-center`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
