---
id: 6858f90c54647a48ddb8822d
title: Passo 11
challengeType: 0
dashedName: step-11
---

# --description--

Neste ponto, o conteúdo, ou seja, os dois elementos `div` dentro do `div` externo serão empilhados um sobre o outro independentemente do tamanho da tela.

Isso acontece porque você tem a classe `flex-col` no elemento `div` que envolve.Para corrigir isso, use o prefixo `lg:` para adicionar a classe `flex-row` ao `div`. 

Isso significa que em dispositivos com telas maiores, a direção do flex será `row`.Para ver isso em efeito novamente, ajuste o navegador para uma largura maior.

# --hints--

Seu elemento externo `div` deve ter a classe utilitária `lg:flex-row`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("lg:flex-row"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
