---
id: 6858eadd3144c89d45e02bc7
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Se você perceber, o `div` externo está muito colado ao topo e o conteúdo dentro dele está muito amontoado à esquerda. Então, você deve trabalhar no espaçamento agora.

Adicione a classe `mt-8` ao `div` para empurrá-lo um pouco para baixo e `p-4` para afastar o texto da esquerda.

# --hints--

Seu elemento externo `div` deve ter a classe utilitária `mt-8`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mt-8"));
```

Seu elemento externo `div` deve ter a classe utilitária `p-4`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
