---
id: 6858ebf2552046ab3de9f17c
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Neste ponto, você notará que há muito espaço à direita do conteúdo no `div`. Você pode minimizar isso adicionando uma largura.

Adicione a classe `md:w-1/2` ao `div` para reduzir a largura para 50% em telas médias e maiores. Ou seja, telas com 768 pixels ou mais.

Para ver isso em ação, expanda a aba de visualização.

# --hints--

Seu elemento externo `div` deve ter a classe utilitária `md:w-1/2`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("md:w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
