---
id: 68e45a2b4ab9b2f48b1cc6df
title: Layoutseite für Preispläne entwerfen
challengeType: 25
dashedName: design-a-pricing-plans-layout-page
demoType: onClick
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Ihre Seite sollte ein `h1`-Element mit dem Text `Pricing Plans` enthalten.
2. Ihre Seite sollte ein `div`-Element mit der Klasse `pricing-container` unterhalb des `h1`-Elements haben.
    - Der `.pricing-container`-Selektor sollte eine `display`-Eigenschaft mit dem Wert `flex` und eine `flex-wrap`-Eigenschaft mit dem Wert `wrap` besitzen.
3. Innerhalb des `.pricing-container`-Elements sollten Sie drei `div`-Elemente mit der Klasse `pricing-card` haben, die die Preispläne darstellen.
    - Eines der `.pricing-card`-Elemente sollte zusätzlich zur `basic-plan`-Klasse die Klasse `pricing-card` haben.
    - Eines der `.pricing-card`-Elemente sollte zusätzlich zur `pro-plan`-Klasse die Klasse `pricing-card` haben.
    - Eines der `.pricing-card`-Elemente sollte zusätzlich zur `premium-plan`-Klasse die Klasse `pricing-card` haben.
4. Ihr `.basic-plan`-Element sollte ein `h2`-Element mit dem Text `Basic` enthalten.
5. Ihr `.basic-plan`-Element sollte ein `p`-Element mit dem Text `$9/month` enthalten.
6. Ihr `.pro-plan`-Element sollte ein `h2`-Element mit dem Text `Pro` enthalten.
7. Ihr `.pro-plan`-Element sollte ein `p`-Element mit dem Text `$19/month` enthalten.
8. Ihr `.premium-plan`-Element sollte ein `h2`-Element mit dem Text `Premium` enthalten.
9. Ihr `.premium-plan`-Element sollte ein `p`-Element mit dem Text `$29/month` enthalten.
10. Jedes Ihrer `.pricing-card`-Elemente sollte:
    - Flexbox verwenden, um seinen Inhalt in einer Spalte anzuordnen und den Abstand zwischen den Kindern mit `space-between` zu verteilen.
    - Die `flex`-Eigenschaft auf `0 0 200px` setzen, um eine einheitliche Breite zu gewährleisten und zu verhindern, dass es im Layout wächst oder schrumpft.
    - Die `border`-Eigenschaft auf `2px solid black` setzen, um zu sehen, wie die verschiedenen Karten Platz einnehmen.
11. Das `.basic-plan`-Element sollte als erstes im Layout erscheinen. Verwenden Sie dafür die `order`-Eigenschaft.
12. Das `.pro-plan`-Element sollte als zweites im Layout erscheinen. Verwenden Sie die `order`-Eigenschaft und setzen Sie dessen `flex-grow`-Eigenschaft auf `2`, damit es mehr Platz als die anderen Pläne einnimmt.
13. Das `.premium-plan`-Element sollte als letztes im Layout erscheinen. Verwenden Sie dafür die `order`-Eigenschaft.

**Hinweis:** Vergewissern Sie sich, dass Sie Ihr Stylesheet in Ihrem HTML verlinken und Ihr CSS anwenden.

# --hints--

Sie sollten ein `h1`-Element haben.

```js
assert.exists(document.querySelector('h1'));
```

Ihre Seite sollte ein `h1`-Element mit dem Text `Pricing Plans` enthalten.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Pricing Plans');
```

Sie sollten ein `div`-Element unterhalb des `h1`-Elements haben.

```js
assert.exists(document.querySelector('h1 + div'));
```

Ihr `div`-Element sollte die Klasse `pricing-container` haben.

```js
assert.exists(document.querySelector('h1 + div.pricing-container'));
```

Sie sollten einen `.pricing-container`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-container'));
```

Ihr `pricing-container`-Selektor sollte die Eigenschaft `display` mit dem Wert `flex` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.display, 'flex');
```

Ihr `.pricing-container`-Selektor sollte die Eigenschaft `flex-wrap` mit dem Wert `wrap` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.flexWrap, 'wrap');
```

Sie sollten drei `div`-Elemente mit der Klasse `pricing-card` innerhalb des `.pricing-container`-Elements haben.

```js
const divEls = document.querySelectorAll('.pricing-container .pricing-card');
assert.lengthOf(divEls, 3);
```

Eines Ihrer `.pricing-card`-Elemente sollte die Klasse `basic-plan` haben.

```js
assert.exists(document.querySelector('div.pricing-card.basic-plan'));
```

Eines Ihrer `.pricing-card`-Elemente sollte die Klasse `premium-plan` haben.

```js
assert.exists(document.querySelector('div.pricing-card.premium-plan'));
```

Eines Ihrer `.pricing-card`-Elemente sollte die Klasse `pro-plan` haben.

```js
assert.exists(document.querySelector('div.pricing-card.pro-plan'));
```

Sie sollten einen `.pricing-card`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-card'));
```

Ihr `.pricing-card`-Selektor sollte die Eigenschaft `display` mit dem Wert `flex` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.display, 'flex');
```

Ihr `.pricing-card`-Selektor sollte die Eigenschaft `flex-direction` mit dem Wert `column` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flexDirection, 'column');
```

Ihr `.pricing-card`-Selektor sollte die Eigenschaft `justify-content` mit dem Wert `space-between` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.justifyContent, 'space-between');
```

Ihr `.pricing-card`-Selektor sollte die Eigenschaft `flex` mit dem Wert `0 0 200px` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flex, '0 0 200px');
```

Ihr `.pricing-card`-Selektor sollte die Eigenschaft `border` mit dem Wert `2px solid black` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.border, '2px solid black');
```

Ihr `.basic-plan`-Selektor sollte die Eigenschaft `order` mit dem Wert `0` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.basic-plan')?.order, '0');
```

Ihr `.basic-plan`-Element sollte ein `h2`-Element mit dem Text `Basic` enthalten.

```js
const h2 = document.querySelector('.basic-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Basic');
```

Ihr `.basic-plan`-Element sollte ein `p`-Element mit dem Text `$9/month` enthalten.

```js
const p = document.querySelector('.basic-plan > p');
assert.strictEqual(p?.textContent.trim(), '$9/month');
```

Ihr `.pro-plan`-Element sollte ein `h2`-Element mit dem Text `Pro` enthalten.

```js
const h2 = document.querySelector('.pro-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Pro');
```

Ihr `.pro-plan`-Element sollte ein `p`-Element mit dem Text `$19/month` enthalten.

```js
const p = document.querySelector('.pro-plan > p');
assert.strictEqual(p?.textContent.trim(), '$19/month');
```

Ihr `.pro-plan`-Element sollte die Eigenschaft `order` mit dem Wert `1` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.order, '1');
```

Ihr `.pro-plan`-Element sollte die Eigenschaft `flex-grow` mit dem Wert `2` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.flexGrow, '2');
```

Ihr `.premium-plan`-Element sollte ein `h2`-Element mit dem Text `Premium` enthalten.

```js
const h2 = document.querySelector('.premium-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Premium');
```

Ihr `.premium-plan`-Element sollte ein `p`-Element mit dem Text `$29/month` enthalten.

```js
const p = document.querySelector('.premium-plan > p');
assert.strictEqual(p?.textContent.trim(), '$29/month');
```

Ihr `premium-plan`-Element sollte das letzte Element im Layout sein und die Eigenschaft `order` mit dem Wert `2` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.premium-plan')?.order, '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Pricing Plans Layout Page</title>
  </head>
  <body></body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Pricing Plans Layout</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Pricing Plans</h1>
    <div class="pricing-container">
      <div class="pricing-card pro-plan">
        <h2>Pro</h2>
        <p class="price">$19/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card basic-plan">
        <h2>Basic</h2>
        <p class="price">$9/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="cross">✗</span> Advanced analytics</li>
          <li><span class="cross">✗</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card premium-plan">
        <h2>Premium</h2>
        <p class="price">$29/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="check">✓</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  align-items: center;
  justify-content: center;
  flex-direction: column;
  min-height: 100vh;
}

h1 {
  margin-bottom: 20px;
}

h2,
p {
  margin: 0;
}

.pricing-container {
  display: flex;
  flex-wrap: wrap;
  align-content: center;
  border: 2px solid #000;
  gap: 20px;
  width: 80%;
  max-width: 900px;
  height: 70vh;
  padding: 0 20px;
}

.pricing-card {
  background: rgb(247, 247, 247);
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
  padding: 20px;
  flex: 0 0 200px;
  text-align: center;
  transition: transform 0.3s ease;
  display: flex;
  flex-direction: column;
  justify-content: space-between;
  border: 2px solid black;
}

.pricing-card:hover {
  transform: scale(1.05);
}

.price {
  font-size: 1.5rem;
  font-weight: bold;
  color: #3e3e3e;
  margin: 10px;
  border-bottom: 1px solid #433d28;
}

.features {
  list-style-type: none;
  padding: 0;
  text-align: left;
}

.features li {
  padding: 5px 0;
  font-size: 0.95rem;
}

.check {
  color: #219c54;
  font-weight: bold;
  margin-right: 8px;
}

.cross {
  color: crimson;
  font-weight: bold;
  margin-right: 8px;
}

a {
  text-decoration: none;
  display: inline-block;
  margin-top: 15px;
  background: #f1be32;
  color: #000;
  font-weight: bold;
  border: 3px solid #eebe3a;
  padding: 5px;
  transition: background 0.3s;
}

a:hover {
  background: #f7c436;
}

.basic-plan {
  order: 0;
}
.pro-plan {
  order: 1;
  flex-grow: 2;
  background: #f0f8ff;
}
.premium-plan {
  order: 2;
}

@media screen and (max-width: 1000px) {
  .pricing-container {
    height: auto;
    padding-bottom: 20px;
  }

  h1 {
    display: block;
    font-size: 1.5rem;
    text-align: center;
  }

  .pricing-card {
    flex: 0 0 175px;
  }
}

@media screen and (max-width: 650px) {
  .pricing-container {
    flex-direction: column;
    height: auto;
  }

  .pricing-card {
    flex: 1 1 auto;
    width: 100%;
  }
}
```

