---
id: 66bb3e20d3dc5b6d0a21f5dd
title: Construye un Contador en Tiempo Real
challengeType: 25
dashedName: build-a-real-time-counter
demoType: onClick
---

# --description--

En este laboratorio, construirás un contador de caracteres en tiempo real utilizando JavaScript. El contador de caracteres mostrará el número de caracteres ingresados en un elemento `textarea`. El contador se actualizará en tiempo real a medida que el usuario escriba en el `textarea`.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**Historia de usuario:**

1. Debe haber un elemento `textarea` con el `id` de `text-input`.
2. Debe haber un elemento `p` con el `id` de `char-count`. Inicialmente debe contener el texto `Character Count: 0/50`. Este texto de marcador de posición debe ser reemplazado con el conteo actual de caracteres en el elemento `textarea`.
3. Cuando el elemento `#text-input` contiene el texto `hello coder`, el elemento `#char-count` debe contener el texto `"Character Count: 11/50"`
4. Cuando el elemento `#text-input` contiene el texto `i am a programmer`, el elemento `#char-count` debe contener el texto `"Character Count: 17/50"`
5. Cuando el elemento `#text-input` contiene el texto `hello world`, el elemento `#char-count` debe contener el texto `"Character Count: 11/50"`
6. Cuando el elemento `#text-input` contiene el texto `I am learning a new language and it is called c++.`, el elemento `#char-count` debe contener el texto `"Character Count: 50/50"`. **NOTA**: Aunque el atributo `maxlength` conseguiría la función correcta, este laboratorio requiere una solución solo con JavaScript.
7. El usuario no debería poder ingresar más de `50` caracteres en el elemento `textarea`. When the character count reaches `50`, any extra input should be automatically trimmed and the text `Character Count: 50/50` should be displayed in `red`.

**Nota:** Asegúrate de enlazar tu hoja de estilos y el archivo JavaScript en tu HTML.

# --hints--

Debe haber un elemento `textarea` con el `id` de `text-input`.

```js
const textInput = document.querySelector('textarea');
assert.equal(textInput.id, 'text-input');
```

Debe haber un elemento `p` con el `id` de `char-count`.

```js
const charCount = document.querySelector('p');
assert.equal(charCount.id, 'char-count');
```

El elemento `#char-count` debería contener inicialmente el texto `Character Count: 0/50`.

```js
const charCount = document.querySelector('p')?.innerText;
assert.equal(charCount, 'Character Count: 0/50');
```

Cuando el elemento `#text-input` contiene el texto `hello coder`, el elemento `#char-count` debe contener el texto `"Character Count: 11/50"`

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "hello coder";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(charCount.innerText.trim(), 'Character Count: 11/50');
```

Cuando el conteo de caracteres es `50`, el texto debe mostrarse en `rojo`.

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "I am learning a new language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
const computedStyle = window.getComputedStyle(charCount);
assert.oneOf(computedStyle.color, [
  "red",
  "rgb(255, 0, 0)",
  "rgb(100%, 0%, 0%)",
  "rgba(255, 0, 0, 1)",
  "rgba(100%, 0%, 0%, 1)",
  "hsl(0, 100%, 50%)",
  "hsla(0, 100%, 50%, 1)"
]);
```

If character count is greater than or equal to `50`, the extra input should be trimmed.

```js
const textInput = document.getElementById('text-input');
textInput.value = "I am learning a new programming language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(textInput.value.length, 50);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real Time Counter</title>
</head>

<body>

</body>


</html>
```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real-time Character Counter</title>
    <link rel="stylesheet" href="styles.css">

<body>
    <div class="container">
        <h1>Real-Time Character Counter</h1>
        <textarea id="text-input" placeholder="Type something..."></textarea>
        <p id="char-count">Character Count: 0/50</p>
    </div>

<script src="script.js"></script>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    font-family: Arial, sans-serif;
    background-color: #c0c1ec;
    margin: 0;
}

h1{
    background-color: #f2f48b;
    padding-left: 10px;
    padding-right: 10px;
}
.container {
    text-align: center;
}

textarea {
    width: 300px;
    height: 100px;
    font-size: 18px;
    padding: 10px;
    border: 1px solid #ccc;
    border-radius: 5px;
    resize: none;
    display: block;
    margin: 0 auto;
    border: 1px solid slategray;
    box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

textarea:focus {
    outline: none;
    border-color: #f2f48b;
}

#char-count {
    margin-top: 10px;
    font-size: 18px;
    color: darkslategray;
    font-weight: bold;
}
```

```js
document.getElementById("text-input").addEventListener("input", function () {
  let textInput = document.getElementById("text-input");
  let charCount = textInput.value.length;

  if (charCount > 50) {
    textInput.value = textInput.value.substring(0, 50);
    charCount = 50;
  }

  let charCountDisplay = document.getElementById("char-count");
  charCountDisplay.innerHTML = `Character Count: ${charCount}/50`;

  if (charCount === 50) {
    charCountDisplay.style.color = "red";
  } 

});
```
