---
id: 60a3e3396c7b40068ad69971
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

该写 CSS 了。

Even though your `<div>` has no text, it's still treated as a box with content. 编写一个使用 `.canvas` 类选择器的 CSS 规则，并将其 `width` 设置为 500 像素。 下面是一个将类 `card` 的宽度设置为 300 像素的 CSS 规则：

```css
.card {
  width: 300px;
}
```

# --hints--

你的代码应该有一个 `.canvas` 选择器。

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

你应该将 `width` 属性设置为 `500px`。

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

你的 `.canvas` 元素应该有一个 `width`，值为 `500px`。

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
