---
id: 60a3e3396c7b40068ad6996e
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

外边距是盒子外的区域，可用于控制与其他盒子或元素之间的空间。

在这里，底部元素的上外边距更大，将其推向页面的更下方。

现在你了解了 CSS 盒子模型，让我们开始 Rothko 绘画。

移除 `<img>` 元素。

# --hints--

你的代码中不应有 `img` 元素。

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
