---
id: 60a3e3396c7b40068ad69977
title: 步骤14
challengeType: 0
dashedName: step-14
---

# --description--

用内边距调整元素内的空间。

在 `.frame` 中，用 `padding` 简写属性将 `.frame` 和 `.canvas` 元素间距增加到 `50px`。 该简写属性将增加元素边框和画布内部的上下左右的空间。

# --hints--

你应将 `padding` 属性设置为 `50px`。

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

你的 `.frame` 元素应有一个 `padding`，值为 `50px`。

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
