---
id: 68e97fe79367ad7b5dd6c9cd
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

在此步骤中，为你的页面添加一个带有文本 `XYZ Bookstore` 的 `h1` 元素。

# --hints--

你应该有一个 `h1` 元素。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素的文本应为 `XYZ Bookstore`。请仔细查看拼写或大小写错误。

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'XYZ Bookstore');
```

你的 `h1` 元素应位于 `body` 元素内。

```js
assert.match(code, /<body>[\s\S]*<h1>[\s\S]*<\/h1>[\s\S]*<\/body>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>
--fcc-editable-region--
<body>

</body>
--fcc-editable-region--
</html>
```
