---
id: 68ebdbacdd3fa474132cc975
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

在 `h1` 元素下方，添加一个带有此文本的 `p` 元素：`Browse our collection of amazing books!`。

# --hints--

你应该有一个 `p` 元素。

```js
assert.exists(document.querySelector('p'));
```

你的 `p` 元素的文本应为：`Browse our collection of amazing books!`。 

```js
assert.equal(document.querySelector('p')?.innerText.trim(), 'Browse our collection of amazing books!');
```

你的 `p` 元素应该在 `h1` 元素的下面。

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  --fcc-editable-region--

  --fcc-editable-region--
</body>

</html>
```
