---
id: 68eca4c4b0f952cf09fabe09
title: 步骤15
challengeType: 0
dashedName: step-15
---

# --description--

在第二个具有 `card` 类的元素内，添加一个带有文本 `Dave's Cooking Adventure` 的 `h2` 元素。

# --hints--

你应该在第二个具有 `card` 类的元素内嵌套一个 `h2` 元素。

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2'));
```

第二个具有 `card` 类的元素内的 `h2` 元素应具有 `Dave's Cooking Adventure` 的文本。

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('h2')?.innerText.trim(), "Dave's Cooking Adventure");
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
  --fcc-editable-region--
    <div class="card" id="dave-cooking-book">

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
