---
id: 68eca2a795b333ca5fee30a8
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

现在创建第二个书籍卡片。为你的页面添加另一个 `div` 元素，`class` 属性设置为 `card`。注意你可以为多个元素重用相同的类名以应用一致的样式。

# --hints--

你应该有两个带有 `class` 属性且值为 `card` 的 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('.card'), 2);
```

第二个具有 `card` 类的元素应位于具有 `card-container` 类的元素内。

```js
assert.lengthOf(document.querySelectorAll('.card-container .card'), 2);
```

具有 `card` 类的两个元素应该是兄弟元素。

```js
assert.exists(document.querySelector('.card + .card'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
  --fcc-editable-region--
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>

  --fcc-editable-region--
  </div>
</body>

</html>
```
