---
id: 68ec6d9a315221aa31e54816
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

`class` 属性用于标识一个或多个元素以进行样式设置。与 `id` 属性不同，类名不需要唯一：多个元素可以共享相同的类。

这是一个示例：

```html
<p class="example">example paragraph</p>
```

为你的 `div` 元素添加一个 `class` 属性，并将其值设置为 `card-container`。

# --hints--

你的 `div` 元素应具有 `class` 属性。

```js
assert.isTrue(document.querySelector('div')?.hasAttribute('class'));
```

你的 `div` 元素应具有值为 `card-container` 的 `class` 属性。

```js
assert.equal(document.querySelector('div')?.className, 'card-container');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
--fcc-editable-region-- 
  <div>

  </div>
--fcc-editable-region--
</body>

</html>
```
