---
id: 68ec99e478211dc578699944
title: 步骤11
challengeType: 0
dashedName: step-11
---

# --description--

在第一个具有 `card` 类的元素中 `h2` 元素下方，添加一个带有以下文本的 `p` 元素：

```md
This is an epic story of Sally and her dog Rex as they navigate through other worlds.
```

# --hints--

你应该在具有 `card` 类的元素内部嵌套一个 `p` 元素。

```js
assert.exists(document.querySelector('.card p'));
```

你的 `p` 元素的文本应为 `This is an epic story of Sally and her dog Rex as they navigate through other worlds.`

```js
assert.equal(document.querySelector('.card p')?.innerText.trim(), "This is an epic story of Sally and her dog Rex as they navigate through other worlds.");
```

你的 `p` 元素应该在 `h2` 元素的下方。

```js
assert.exists(document.querySelector('.card h2 + p'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
  --fcc-editable-region--
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
