---
id: 68ec9332a9b5b2b32487bd00
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

`id` 属性为超文本标记语言元素添加唯一标识符。每个 `id` 应在页面内唯一且只能使用一次。

`id` 值不能包含空格，只能包含字母、数字、下划线和连字符。

这是一个示例：

```html
<p id="para">example paragraph</p>
```

为你的类为 `card` 的元素添加一个 `id` 属性，并将其值设置为 `sally-adventure-book`。

# --hints--

你的具有 `card` 类的元素应具有 `id` 属性。

```js
assert.isTrue(document.querySelector('.card')?.hasAttribute('id'));
```

你的元素具有 `card` 类时，应具有值为 `sally-adventure-book` 的 `id`。

```js
assert.equal(document.querySelector('.card')?.id, 'sally-adventure-book');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
  --fcc-editable-region--
    <div class="card">

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
