---
id: 68eca11e99e3c5c894ca9d69
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

`button` 元素用于在网页上创建可点击的按钮。按钮是用户可以点击以执行操作的交互式元素。

在具有 `class` 为 `card` 的元素内添加一个 `button` 元素，为该按钮添加一个 `class` 属性并设置为 `btn`，以及文本 `Buy Now`。

# --hints--

你应该在具有 `card` 类的元素内部嵌套一个 `button` 元素。

```js
assert.exists(document.querySelector('.card button'));
```

你的 `button` 元素应具有 `class` 属性。

```js
assert.isTrue(document.querySelector('.card button')?.hasAttribute('class'));
```

你的 `button` 元素应有一个 `class`，其值为 `btn`。

```js
assert.equal(document.querySelector('.card button')?.className, 'btn');
```

你的 `button` 元素的文本应为 `Buy Now`。 

```js
assert.equal(document.querySelector('.card button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
  --fcc-editable-region--
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
