---
id: 68eca5c0065c82d256d29ca3
title: 第 17 步
challengeType: 0
dashedName: step-17
---

# --description--

在第二个卡片内，添加一个 `button` 元素，其 `class` 属性设置为 `btn`，文本为 `Buy Now`。

两个 `button` 元素现在共享相同的 `class`，这意味着它们可以一起进行一致的样式设置。

# --hints--

你应该在第二个具有 `card` 类的元素内嵌套一个 `button` 元素。

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('button'));
```

你的第二张卡片的 `button` 元素应具有 `class` 属性。

```js
const cards = document.querySelectorAll('.card');
assert.isTrue(cards[1]?.querySelector('button')?.hasAttribute('class'));
```

你的第二张卡片的 `button` 元素应具有值为 `btn` 的 `class`。

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.className, 'btn');
```

你第二张卡片的 `button` 元素的文本应为 `Buy Now`。

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
  --fcc-editable-region--
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
  
    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
