---
id: 68eca6e5a6759ed4ea0034dc
title: 步骤 19
challengeType: 0
dashedName: step-19
---

# --description--

在具有 `btn-container` 类的元素内，添加两个 `button` 元素：

第一个按钮：

- ID：`view-cart-btn`
- 类：`btn`
- 文本：`View Cart`

第二个按钮：

- 编号：`checkout-btn`
- 类：`btn`
- 文本：`Checkout`


恭喜！你已成功使用 div、类 和 id 构建了书店页面的结构来组织你的内容。

# --hints--

你应该在 `.btn-container` 元素内嵌套两个 `button` 元素。

```js
assert.lengthOf(document.querySelectorAll('.btn-container button'), 2);
```

在类为 `btn-container` 的元素内，你的第一个按钮应具有值为 `view-cart-btn` 的 `id` 属性。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.id, 'view-cart-btn');
```

在具有 `btn-container` 类的元素内，你的第一个按钮应具有值为 `btn` 的 `class` 属性。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.className, 'btn');
```

你第一个按钮的文本应为 `View Cart`。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.innerText.trim(), 'View Cart');
```

具有 `btn-container` 类的元素内的第二个按钮应具有值为 `checkout-btn` 的 `id` 属性。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.id, 'checkout-btn');
```

位于具有 `class="btn-container"` 的元素内的第二个按钮应具有值为 `btn` 的 `class` 属性。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.className, 'btn');
```

你第二个按钮的文本应为 `Checkout`。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.innerText.trim(), 'Checkout');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
--fcc-editable-region--
  <p>Review your selections and continue to checkout.</p>
  <div class="btn-container">

  </div>
--fcc-editable-region--
</body>

</html>
```

# --solutions--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>

  <div class="card-container">
     <div class="card" id="sally-adventure-book">
     <h2>Sally's SciFi Adventure</h2>
     <p>
          This is an epic story of Sally and her dog Rex as they navigate
          through other worlds.
     </p>
     <button class="btn">Buy Now</button>
     </div>

     <div class="card" id="dave-cooking-book">
     <h2>Dave's Cooking Adventure</h2>
     <p>
          This is the story of Dave as he learns to cook everything from
          pancakes to pasta, one recipe at a time.
     </p>
     <button class="btn">Buy Now</button>
     </div>
  </div>

  <p>Review your selections and continue to checkout.</p>

  <div class="btn-container">
     <button id="view-cart-btn" class="btn">View Cart</button>
     <button id="checkout-btn" class="btn">Checkout</button>
  </div>
</body>

</html>
```
