---
id: 68ec6cd5b7e8f5a8f7319e32
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

`div` 元素用作容器来分组其他超文本标记语言元素。当你想分组将共享一组 CSS 样式的超文本标记语言元素时，主要会使用 `div` 元素。

在 `p` 元素下方，添加一个 `div` 元素。这个 `div` 将作为你的书籍卡片的容器。

# --hints--

你应该有一个 `div` 元素。

```js
assert.exists(document.querySelector('div'));
```

你的 `div` 元素应位于你的 `p` 元素下方。

```js
assert.exists(document.querySelector('p + div'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
--fcc-editable-region--

--fcc-editable-region--
</body>

</html>
```
