---
id: 68ec98b38d83a3c28dd30efe
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

在第一个具有 `card` 类的元素内，添加一个带有文本 `Sally's SciFi Adventure` 的 `h2` 元素。

# --hints--

你应该在具有 `card` 类的元素内嵌套一个 `h2` 元素。

```js
assert.exists(document.querySelector('.card h2'));
```

你的 `h2` 元素的文本应为 `Sally's SciFi Adventure`。 

```js
assert.equal(document.querySelector('.card h2')?.innerText.trim(), "Sally's SciFi Adventure");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
  --fcc-editable-region--
    <div class="card" id="sally-adventure-book">

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
