---
id: 68eca6364c5616d393389a30
title: 步骤18
challengeType: 0
dashedName: step-18
---

# --description--

请记住，超文本标记语言元素看起来是这样的：

```html
<element attribute="value">
    inner text
</element>
```

在类为 `card-container` 的元素下方，添加一个带有此文本的新的 `p` 元素：

```md
Review your selections and continue to checkout.
```

在 `p` 元素下方，创建一个 `div` 元素，其 `class` 属性设置为 `btn-container`。该容器将分组你的导航按钮元素。

# --hints--

新添加的 `p` 元素应位于具有 `card-container` 类的元素下方。

```js
assert.exists(document.querySelector('.card-container + p'));
```

你的新建 `p` 元素的文本应为 `Review your selections and continue to checkout.`

```js
const cardContainer = document.querySelector('.card-container');
assert.equal(cardContainer?.nextElementSibling?.innerText.trim(), 'Review your selections and continue to checkout.');
```

你应该有一个 `div` 元素，其 `class` 为 `btn-container`。

```js
assert.exists(document.querySelector('div.btn-container'));
```

你的 `div` 元素具有 `btn-container` 的 `class` 应该位于新添加的 `p` 元素之后。

```js
assert.exists(document.querySelector('p + .btn-container'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
--fcc-editable-region--
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>

--fcc-editable-region--
</body>

</html>
```
