---
id: 69134b6cb43aa5f254950a10
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次研讨会中，你将通过创建显示不同书籍信息的书籍卡片来构建一个书店页面。你将练习使用 `div` 元素、类和 ID 来组织内容。

通过创建超文本标记语言 公式化 来开始你的书店页面。

为你的文档添加 `<!DOCTYPE html>` 声明和 `html` 及 `head` 元素。

为你的 `html` 元素添加一个 `lang` 属性并将其设置为 `"en"`。

# --hints--

你的代码应包含一个 `<!DOCTYPE html>` 声明。

```js
assert.match(code, /<!DOCTYPE html>/i);
```

你的代码应包含开闭的 `<html>` 标签。

```js
assert.match(code,/<html\b[^>]*>[\s\S]*<\/html>/i);
```

你的 `html` 元素应具有设置为 `"en"` 的 `lang` 属性。

```js
assert.match(code, /<html[^>]*\blang\s*=\s*["']en["'][^>]*>/i);
```

你的代码应包含开头和结尾的 `<head>` 标签。

```js
assert.match(code,/<head>[\s\S]*<\/head>/i);
```

你的 `head` 元素应嵌套在 `html` 元素内。

```js
assert.match(code, /<\/html>\s*$/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
