---
id: 69134d14150d51f8066a3246
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

现在，改进你的 HTML 文档结构以确保你的页面编码正确。

在 `head` 元素内，添加 `<meta charset="UTF-8">` 元素。

最后，在 `head` 部分下方添加一个 `body` 元素。这是所有可见页面内容所在的位置。

# --hints--

你的代码应包含一个 `<meta charset="UTF-8">` 元素。

```js
assert.exists(document.querySelector('head > meta[charset="utf-8"]'));
```

你的代码应包含开闭的 `<body>` 标签。

```js
assert.match(code,/<body>[\s\S]*<\/body>/i);
```

你的 `body` 元素应位于 `head` 元素下方。

```js
assert.match(code, /<\/head>[\s\S]*<body>[\s\S]*<\/body>[\s\S]*<\/html>/i);
```

# --seed--
## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html lang="en">
<head>
  <title>XYZ Bookstore Page</title>
</head>

</html>
--fcc-editable-region--
```
