---
id: 691f85085787a8c4ca589fbb
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

在 `head` 元素内添加 `title` 元素。 

将页面标题设置为 `XYZ Bookstore Page`。

# --hints--

你的代码应包含开头和结尾的 `<title>` 标签。 

```js
assert.match(code,/<title>[\s\S]*<\/title>/i);
```

你的标题应为 `XYZ Bookstore Page`。

```js
assert.equal(document.querySelector('title')?.innerText.trim(),'XYZ Bookstore Page');
```

你的 `title` 元素应嵌套在 `head` 元素内。

```js
assert.match(code, /<head>[\s\S]*<title>[\s\S]*<\/title>[\s\S]*<\/head>/i);
```

# --seed--

## --seed-contents--

```html
  --fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
<head>

</head>
  --fcc-editable-region-- 
</html>
```
