---
id: 587d824a367417b2b2512c46
title: Aprende cómo funcionan las aserciones de JavaScript
challengeType: 2
forumTopicId: 301589
dashedName: learn-how-javascript-assertions-work
---

# --description--

Trabajar en estos desafíos implica escribir tu código usando uno de los siguientes métodos:

- Clone <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">this GitHub repo</a> and complete these challenges locally.
- Use a site builder of your choice to complete the project. Be sure to incorporate all the files from our GitHub repo.

# --instructions--

Dentro de `tests/1_unit-tests.js`, bajo la prueba etiquetada como `#1` en el conjunto de prueba `Basic Assertions`, cambia cada `assert` a `assert.isNull` o `assert.isNotNull` para pasar la prueba (deber evaluarse como `true`). No modifiques los argumentos pasados ​​a los verificadores.

# --hints--

Todas los test deben pasar.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Debes elegir el método correcto para la primera aserción - `isNull` vs. `isNotNull`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isNull', 'Null is null');
```

Debes elegir el método correcto para la segunda aserción - `isNull` vs. `isNotNull`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'isNotNull', '1 is not null');
```

