---
id: 587d824e367417b2b2512c57
title: Comprueba si un objeto es una instancia de un constructor
challengeType: 2
forumTopicId: 301605
dashedName: test-if-an-object-is-an-instance-of-a-constructor
---

# --description--

Como recordatorio, este proyecto se está construyendo sobre la siguiente estructura inicial clonada desde <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`#instanceOf` verifica que un objeto es una instancia de un constructor.

# --instructions--

Dentro de `tests/1_unit-tests.js` bajo la prueba etiquetada `#18` en la suite `Objects`, cambia cada `assert` por `assert.instanceOf` o por `assert.notInstanceOf` para pasar la prueba (debe evaluarse como `true`). No modifiques los argumentos pasados ​​a los verificadores.

# --hints--

Todas los test deben pasar.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Debe elegir el método correcto para la primera aserción - `instanceOf` vs `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notInstanceOf',
  'myCar is not an instance of Plane'
);
```

Debe elegir el método correcto para la segunda aserción - `instanceOf` vs `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'instanceOf',
  'airlinePlane is an instance of Plane'
);
```

Debe elegir el método correcto para la tercera aserción - `instanceOf` vs `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[2].method,
  'instanceOf',
  'everything is an Object in JavaScript...'
);
```

Debe elegir el método correcto para la cuarta aserción - `instanceOf` vs `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'notInstanceOf',
  'myCar.wheels is not an instance of String'
);
```

