---
id: 587d824b367417b2b2512c47
title: Comprueba si una variable o función está definida
challengeType: 2
forumTopicId: 301602
dashedName: test-if-a-variable-or-function-is-defined
---

# --description--

Como recordatorio, este proyecto se está construyendo sobre el siguiente proyecto inicial clonado de <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

Dentro de `tests/1_unit-tests.js` bajo el test etiquetado `#2` en el suite `Basic Assertions`, cambia cada `assert` a `assert.isDefined()` o `assert.isUndefined()` para hacer que el test pase (debe evaluarse como `true`). No modifiques los argumentos pasados ​​a los verificadores.

# --hints--

Todas los test deben pasar.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Debe elegir el método correcto para la primera aserción - `isDefined` vs `isUndefined`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'isDefined',
    'Null is not undefined'
  );
```

Debe elegir el método correcto para la segunda aserción - `isDefined` vs `isUndefined`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'isUndefined',
    'Undefined is undefined'
  );
```

Debe elegir el método correcto para la tercera aserción - `isDefined` vs `isUndefined`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'isDefined',
    'A string is not undefined'
  );
```

