---
id: 587d824c367417b2b2512c4c
title: Confirmar la igualdad profunda con .deepEqual y .notDeepEqual
challengeType: 2
forumTopicId: 301587
dashedName: assert-deep-equality-with--deepequal-and--notdeepequal
---

# --description--

Como recordatorio, este proyecto se está construyendo sobre la siguiente estructura inicial clonada desde <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`deepEqual()` confirma que dos objetos son iguales en profundidad.

# --instructions--

Dentro de `tests/1_unit-tests.js` bajo la prueba etiquetada `#7` en el conjunto de pruebas `Equality`, cambia cada `assert` a `assert.deepEqual` o `assert.notDeepEqual` para pasar la prueba (debe evaluarse como `true`). No modifiques los argumentos pasados ​​a los verificadores.

# --hints--

Todas los test deben pasar.

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
})();
```

Debe elegir el método correcto para la primera aserción - `deepEqual` vs. `notDeepEqual`.

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'deepEqual',
    'The order of the keys does not matter'
  );
})();
```

Debe elegir el método correcto para la segunda aserción - `deepEqual` vs. `notDeepEqual`.

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'notDeepEqual',
    'The position of elements within an array does matter'
  );
})();
```

