---
id: 5f8884f4c46685731aabfc41
title: Ejecuta pruebas funcionales usando un navegador sin interfaz gráfica II
challengeType: 2
forumTopicId: 301594
dashedName: run-functional-tests-using-a-headless-browser-ii
---

# --description--

Como recordatorio, este proyecto se está construyendo sobre la siguiente estructura inicial clonada desde <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

Dentro `tests/2_functional-tests.js`, en el `'Submit the surname "Vespucci" in the HTML form'` prueba (`// #6`), automatiza lo siguiente:

1.  Rellena el formulario con el apellido `Vespucci`
2.  Presione el botón Enviar

Y dentro del callback de `pressButton`:

1.  Compruebe que el estado es OK `200`
2.  Comprueba que el texto dentro del elemento `span#name` es `'Amerigo'`
3.  Comprueba que el texto dentro del elemento `span#surname` es `'Vespucci'`
4.  Compruebe que existen elemento(s) `span#dates` y su contador es `1`

No olvide eliminar la llamada `assert.fail()`.

# --hints--

Todas los test deben pasar.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Debes asegurarte que la petición del navegador sin interfaz gráfica ha sido exitosa.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'browser.success');
```

Debes comprobar que el texto dentro del elemento `span#name` es `'Amerigo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'browser.text');
assert.match(data.assertions[1].args[0], /('|")span#name\1/);
assert.match(data.assertions[1].args[1], /('|")Amerigo\1/);
```

Debes comprobar que el texto dentro del elemento `span#surname` es `'Vespucci'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'browser.text');
assert.match(data.assertions[2].args[0], /('|")span#surname\1/);
assert.match(data.assertions[2].args[1], /('|")Vespucci\1/);
```

Debe comprobar que el elemento `span#dates` existe y su contador es 1.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'browser.elements');
assert.match(data.assertions[3].args[0], /('|")span#dates\1/);
assert.equal(data.assertions[3].args[1], 1);
```

