---
id: 694a9ff65d4472dd451d8255
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este taller, practicarás cómo usar booleanos y declaraciones condicionales en Python construyendo un calculador de reservas de entradas de cine.

Comienza declarando variables que almacenan información sobre el usuario y la función de la película.

Como aprendiste en lecciones anteriores, las variables se asignan usando el operador de asignación (`=`) de esta manera:

```py
# Assigns the value 10 to variable x
x = 10
```

Crea una variable llamada `base_price` para almacenar el precio base del boleto de cine y asigna su valor a `15`. Crea otra variable llamada `age` para almacenar la edad del usuario y asigna su valor a `21`.

# --hints--

Debes declarar una variable llamada `base_price`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("base_price")`)) })
```

Debes asignar el valor entero `15` a `base_price`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("base_price").is_equivalent("base_price = 15")
`)) })
```

Debes declarar una variable llamada `age`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("age")`)) })
```

Debes asignar el valor entero `21` a `age`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("age").is_equivalent("age = 21")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
