---
id: 694cc1b515976889afd4cded
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Ahora necesitas almacenar algunos valores de cadena sobre el boleto de la película. Puedes almacenar cadenas en Python de esta manera:

```py
name = 'Alice'
```

Crea una variable llamada `seat_type` para almacenar el tipo de asiento que el usuario ha seleccionado y asigna su valor a la cadena `Gold`. Crea otra variable llamada `show_time` para almacenar la hora de la función de la película y asigna su valor a la cadena `Evening`.

Recuerda rodear ambos valores con comillas simples o dobles.

# --hints--

Debes declarar una variable llamada `seat_type`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("seat_type")`)) })
```

Debes asignar la cadena `Gold` a `seat_type`. Recuerda rodear el valor con comillas simples o dobles.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("seat_type").is_equivalent("seat_type = 'Gold'")
`)) })
```

Debes declarar una variable llamada `show_time`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("show_time")`)) })
```

Debes asignar la cadena `Evening` a `show_time`. Recuerda rodear el valor con comillas simples o dobles.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("show_time").is_equivalent("show_time = 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
--fcc-editable-region--

--fcc-editable-region--
```
