---
id: 6958e2377cacbc46a7f7383a
title: Paso 13
challengeType: 20
dashedName: step-13
---

# --description--

Ahora, agrega una cláusula `else` a tu sentencia `if` e imprime `No extra charges will be applied` dentro del cuerpo de `else`. Esto imprimirá el mensaje solo cuando la condición de cargos extra sea falsa.

Luego, debajo de la cláusula `else`, usa la llamada `print()` para mostrar un mensaje que muestre `Extra charges:` seguido del valor actualizado de `extra_charges` y verifica la salida en el terminal.

# --hints--

Debes tener una cláusula `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3].find_conditions()[1].is_empty()`)) })
```

Debes imprimir `No extra charges will be applied` dentro de tu cláusula `else`. Recuerda rodear el mensaje con comillas simples o dobles.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[1].has_call("print('No extra charges will be applied')")
`)) })
```

Debes tener `print('Extra charges:', extra_charges)` debajo de tu sentencia `if...else`.

```js
({ test: () => (runPython(`
if_stmt = """if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')"""
print_call = "print('Extra charges:', extra_charges)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
--fcc-editable-region--

--fcc-editable-region--
```
