---
id: 6958e59845777391b86d7331
title: Paso 18
challengeType: 20
dashedName: step-18
---

# --description--

En Python, una sentencia `if` también puede colocarse dentro del cuerpo de otra sentencia `if`. Esto se llama una sentencia `if` anidada.

Una estructura `if` anidada te permite verificar una condición adicional solo después de que la primera condición ya se haya cumplido. La estructura `if` interna se ejecutará solo si la condición del `if` externo es verdadera.

```py
if condition1:
    # Code to execute if condition1 is True
    if condition2:
        # Code to execute if both conditions are True
```

Ahora calcularás los cargos por servicio según el tipo de asiento que el usuario haya seleccionado.

Todavía dentro del cuerpo de tu última sentencia `if`, crea una variable llamada `service_charges` y asígnale el valor `0`.

Luego, crea una estructura `if` anidada para verificar si `seat_type` es igual a `Premium`. Dentro del cuerpo de la estructura `if`, actualiza el valor de `service_charges` a `5`.

# --hints--

Debes declarar una variable llamada `service_charges` dentro del cuerpo de tu última sentencia `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_variable("service_charges")
`)) })
```

Debes asignar el valor entero `0` a `service_charges`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 0")
`)) })
```

Debes tener una sentencia `if` dentro del cuerpo de tu última sentencia `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_ifs()[0]`)) })
```

La nueva condición de la sentencia `if` debe ser `seat_type == 'Premium'`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("seat_type == 'Premium'")
`)) })
```

Debes asignar el valor entero `5` a la variable `service_charges` dentro de tu nueva estructura `if`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_ifs()[0].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 5")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

if age >= 21 or age >= 18 and (show_time != 'Evening' or is_member):
    print('Ticket booking condition satisfied')

--fcc-editable-region--
    
--fcc-editable-region--
else:
    print('Ticket booking failed due to restrictions')
```
