---
id: 6958e0ecc121b0ae3dd9c5e0
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

También debes manejar el caso cuando el usuario no califique para un descuento de membresía.

Agrega una cláusula `else` a tu declaración `if` e imprime `User does not qualify for membership discount` dentro del cuerpo de `else`.

Puedes imprimir un mensaje seguido de una variable así:

```py
print('Total:', total)
```

También quieres mostrar el valor actualizado de `discount`. Debajo de la sentencia `if...else`, usa la llamada `print()` para mostrar un mensaje que muestre `Discount:` seguido del valor actualizado de `discount`. Luego, verifica la salida en la terminal.

# --hints--

Debes tener una cláusula `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2].find_conditions()[1].is_empty()`)) })
```

Debes imprimir `User does not qualify for membership discount` dentro de tu cláusula `else`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[1].has_call("print('User does not qualify for membership discount')")
`)) })
```

Debes tener `print('Discount:', discount)` debajo de tu sentencia `if...else`.

```js
({ test: () => (runPython(`
if_stmt = """if is_member:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')"""
print_call = "print('Discount:', discount)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
if is_member:
    discount = 3
    print('User qualifies for membership discount')
--fcc-editable-region--

--fcc-editable-region--
```
