---
id: 6958e1875e592b58e8852757
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

Ahora crea una variable llamada `extra_charges` y asígnale el valor `0`. Esto representará cargos adicionales para aplicar al boleto de cine los fines de semana.

Luego, crea una estructura `if` para verificar si `is_weekend` es verdadero. Dentro del cuerpo de la estructura `if`, actualiza el valor de `extra_charges` a `2` e imprime `Extra charges will be applied` en la terminal.

# --hints--

Debes declarar una variable llamada `extra_charges`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("extra_charges")`)) })
```

Debes asignar el valor entero `0` a `extra_charges`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("extra_charges").is_equivalent("extra_charges = 0")
`)) })
```

Deberías tener una cuarta sentencia `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3]`)) })
```

La cuarta condición de tu sentencia `if` debe ser `is_weekend`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_conditions()[0].is_equivalent("is_weekend")
`)) })
```

Debes asignar el valor entero `2` a la variable `extra_charges` dentro de tu nueva estructura `if`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[3].find_bodies()[0].find_variable("extra_charges") 
assert _var.is_equivalent("extra_charges = 2")
`)) })
```

Debes imprimir `Extra charges will be applied` dentro de tu nueva estructura `if`. Recuerda rodear el mensaje con comillas simples o dobles.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[0].has_call("print('Extra charges will be applied')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

--fcc-editable-region--

--fcc-editable-region--
```
