---
id: 6958dd85587e10c99162dd81
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Como aprendiste en lecciones anteriores, en Python, una sentencia `if` puede usarse para ejecutar código dependiendo de si una condición es verdadera.

Una sentencia `if` consiste en la palabra clave `if`, seguida de una condición y dos puntos. El código que se ejecuta cuando la condición es verdadera, que debe estar indentado, se llama el cuerpo de la sentencia `if`.

```py
if condition:
    # Code to execute if condition is True
```

En este paso, verificarás si el usuario es elegible para reservar un boleto de cine según su edad.

Crea una sentencia `if` para verificar si `age` es mayor que `17`. Dentro del cuerpo de la sentencia `if`, imprime `User is eligible to book a ticket`. Esto imprimirá el mensaje solo cuando la edad del usuario sea mayor que `17`.

Recuerda indentar el cuerpo de la sentencia `if` y rodear el mensaje con comillas simples o dobles dentro de la llamada `print()`.

# --hints--

Debes tener una sentencia `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[0]`)) })
```

La condición de la sentencia `if` debe ser `age > 17`.

```js 
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("age > 17") or _cond.is_equivalent("17 < age")
`)) })
```

Debes tener una llamada a `print()` dentro de tu sentencia `if`. Recuerda indentar el cuerpo de la sentencia `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].find_calls("print")[0]
`)) })
```

Debes imprimir `User is eligible to book a ticket` dentro de tu sentencia `if`. Recuerda rodear el mensaje con comillas simples o dobles.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].has_call("print('User is eligible to book a ticket')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

--fcc-editable-region--

--fcc-editable-region--
```
