---
id: 6958de0d773289a89515a6b8
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Ahora verificarás si el usuario tiene permitido reservar un espectáculo nocturno según su edad.

Crea una estructura `if` para verificar si `age` es mayor o igual a `21`. Dentro del cuerpo de la estructura `if`, imprime `User is eligible for Evening shows`.

# --hints--

Debes tener dos sentencias `if` en tu código.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1]`)) })
```

La segunda condición de tu sentencia `if` debe ser `age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[1].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

Debes imprimir `User is eligible for Evening shows` dentro de tu nueva sentencia `if`. Recuerda rodear el mensaje con comillas simples o dobles.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[0].has_call("print('User is eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

--fcc-editable-region--

--fcc-editable-region--
```
