---
id: 6960b308bf41cc39c9b4b438
title: Paso 14
challengeType: 20
dashedName: step-14
---

# --description--

En Python, el operador `or` se usa para verificar si al menos una de dos condiciones es verdadera. Así es como puedes usarlo en una sentencia `if`:

```py
if condition1 or condition2:
    # Code to execute if any condition is True
```

También deben aplicarse cargos adicionales si el espectáculo es por la noche. Usa el operador `or` para combinar la condición existente de tu sentencia `if` con una segunda condición que verifique si `show_time` es igual a la cadena `Evening`.

# --hints--

La última condición de tu sentencia `if` debe ser `is_weekend or show_time == 'Evening'`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_conditions()[0].is_equivalent("is_weekend or show_time == 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
--fcc-editable-region--
if is_weekend:
--fcc-editable-region--
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)
```
