---
id: 696099c1ac72a4e585de1d5b
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Algunas informaciones solo pueden ser verdaderas o falsas. Como has aprendido en lecciones anteriores, esto puede representarse usando valores booleanos.

Crea una variable llamada `is_member` para indicar si el usuario es miembro y asigna su valor a `True`.

Debajo de la variable `is_member` crea otra variable llamada `is_weekend` para indicar si la función de la película es en un fin de semana y asigna su valor a `False`. No pongas el valor entre comillas.

# --hints--

Debes declarar una variable llamada `is_member`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_member")`)) })
```

Debes asignar el valor booleano `True` a `is_member`. No rodees el valor con comillas.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = True")
`)) })
```

Debes declarar una variable llamada `is_weekend`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_weekend")`)) })
```

Debes asignar el valor booleano `False` a `is_weekend`. No rodees el valor con comillas.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_weekend").is_equivalent("is_weekend = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--

--fcc-editable-region--
```
