---
id: 696099abb6acc539b27055a1
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

En el paso anterior, verificaste una condición usando una sentencia `if`. La cláusula `else` se usa para manejar el caso cuando la condición es falsa. Aquí está la sintaxis de una sentencia `if...else`:

```py
if condition:
    # Code to execute if condition is True
else:
    # Code to execute if condition is False
```

Ahora, agrega una cláusula `else` a tu declaración `if` e imprime `User is not eligible for Evening shows` dentro del cuerpo de `else`. Esto se imprimirá solo cuando la condición en la declaración `if` sea falsa.

# --hints--

Debes tener una cláusula `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1].find_conditions()[1].is_empty()`)) })
```

Debes imprimir `User is not eligible for Evening shows` dentro de tu cláusula `else`. Recuerda rodear el mensaje con comillas simples o dobles.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[1].has_call("print('User is not eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
--fcc-editable-region--

--fcc-editable-region--
```
