---
id: 6960ba6e6c95ed4ad896c6f2
title: Paso 17
challengeType: 20
dashedName: step-17
---

# --description--

Hay una excepción más a las reglas de reserva. Los usuarios entre `18` y `21` pueden reservar funciones nocturnas si son miembros.

Cuando se usan múltiples operadores lógicos en una sentencia `if`, las condiciones unidas con `and` se evalúan antes que las condiciones unidas con `or`. Los paréntesis `()` se usan en Python para agrupar condiciones y controlar el orden en que se evalúan.

```py
if condition1 and (condition2 or condition3):
    # Code to execute if conditions evaluate to True
else:
    # Code to execute if conditions evaluate to False
```

Ahora, añade otra condición a tu `if` existente usando el operador `or` para verificar si `is_member` es verdadero. Usa los paréntesis `()` para agrupar las condiciones `show_time != 'Evening'` e `is_member` juntas como se muestra en el ejemplo anterior.

# --hints--

La última condición de tu sentencia `if` debería ser `age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--
if age >= 21 or age >= 18 and show_time != 'Evening':
--fcc-editable-region--
    print('Ticket booking condition satisfied')
else:
    print('Ticket booking failed due to restrictions')
```
