---
id: 6960a848aacf859f7bb15447
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Ahora cambia el valor de la variable `is_member` a `False` ya que el usuario no es miembro.

Después de eso, verás que el valor de `discount` ahora permanece en `0`, porque ambas condiciones deben cumplirse para que se aplique el descuento.

# --hints--

Debes actualizar el valor de la variable `is_member` a `False`. No rodees el valor con comillas.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--
is_member = True
--fcc-editable-region--
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
