---
id: 6960b6a6ea06e27dcc7eef37
title: Paso 15
challengeType: 20
dashedName: step-15
---

# --description--

Ahora verificarás si el usuario cumple con las condiciones para reservar un boleto de cine. Los usuarios con edad `21` o más siempre pueden reservar boletos sin ninguna restricción.

Crea una estructura `if` para verificar si `age` es mayor o igual a `21`. Dentro del cuerpo de la estructura `if`, imprime `Ticket booking condition satisfied` en el terminal.

Luego, agrega una cláusula `else` a tu declaración `if` e imprime `Ticket booking failed due to restrictions` dentro del cuerpo de `else`.

# --hints--

Deberías tener una quinta sentencia `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4]`)) })
```

La quinta condición de tu sentencia `if` debe ser `age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

Debes imprimir `Ticket booking condition satisfied` dentro de tu nueva estructura `if`. Recuerda rodear el mensaje con comillas simples o dobles.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_call("print('Ticket booking condition satisfied')")
`)) })
```

Tu nueva sentencia `if` debe tener una cláusula `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_conditions()[1] == _Node()`)) })
```

Debes imprimir `Ticket booking failed due to restrictions` dentro de tu cláusula `else`. Recuerda rodear el mensaje con comillas simples o dobles.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[1].has_call("print('Ticket booking failed due to restrictions')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--

--fcc-editable-region--
```
