---
id: 6960a5165f7efe52a0c0aa39
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

En Python, el operador `and` se usa para verificar si múltiples condiciones son verdaderas. Aquí tienes cómo puedes usarlo para combinar dos condiciones en una sentencia `if`:

```py
if condition1 and condition2:
    # Code to execute if all conditions are True
```

El descuento por membresía solo debe aplicarse a los miembros si su `age` es mayor o igual a `21`.

Usa el operador `and` para combinar la condición existente de tu sentencia `if` con otra condición que verifique si `age` es mayor o igual a `21`.

# --hints--

La tercera condición de tu sentencia `if` debe ser `is_member and age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[2].find_conditions()[0] 
assert _cond.is_equivalent("is_member and age >= 21") or _cond.is_equivalent("is_member and 21 <= age")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--
if is_member:
--fcc-editable-region--
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
