---
id: 696cb889fc4fdeb337e52864
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

En Python, las condiciones `if` no tienen que comparar valores explícitamente con `True` o `False`. En cambio, se basan en la veracidad de los valores. Los valores verdaderos evalúan a `True` en un contexto booleano, como la condición de una sentencia `if`. Estos incluyen números distintos de cero, cadenas no vacías y el valor booleano `True`.

```py
if value:
   print('value is truthy')
```

Crea una sentencia `if` para verificar si `is_member` es verdadero. Dentro del cuerpo de la sentencia `if`, actualiza el valor de `discount` a `3` e imprime `User qualifies for membership discount` en el terminal.

# --hints--

Deberías tener tres sentencias `if` en tu código.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2]`)) })
```

La tercera condición de tu sentencia `if` debe ser `is_member`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_conditions()[0].is_equivalent("is_member")
`)) })
```

Debes asignar el valor entero `3` a la variable `discount` dentro de tu nueva estructura `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].find_variable("discount").is_equivalent("discount = 3")
`)) })
```

Debes imprimir `User qualifies for membership discount` dentro de tu nueva sentencia `if`. Recuerda rodear el mensaje con comillas simples o dobles.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].has_call("print('User qualifies for membership discount')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--

--fcc-editable-region--
```
