---
id: 5e44414f903586ffb414c950
title: Costruire un progetto di calcolatore di probabilità
challengeType: 23
saveSubmissionToDB: true
forumTopicId: 462364
dashedName: build-a-probability-calculator-project
---

# --description--

Supponiamo che ci sia un cappello contenente 5 palline blu, 4 palline rosse e 2 palline verdi. Qual è la probabilità che un'estrazione casuale di 4 palline contenga almeno 1 pallina rossa e 2 palline verdi? Anche se sarebbe possibile calcolare la probabilità usando la matematica avanzata, un modo più semplice è scrivere un programma che esegua un gran numero di esperimenti per stimare una probabilità approssimativa.

Per questo progetto, scriverai un programma per determinare la probabilità approssimativa di estrarre casualmente determinate palline da un cappello.

Per prima cosa, crea una classe `Hat` in `main.py`. La classe dovrebbe accettare un numero variabile di argomenti che specificano il numero di palline di ogni colore presenti nel cappello. Per esempio, un oggetto della classe potrebbe essere creato in uno qualsiasi di questi modi:

```py
hat1 = Hat(yellow=3, blue=2, green=6)
hat2 = Hat(red=5, orange=4)
hat3 = Hat(red=5, orange=4, black=1, blue=0, pink=2, striped=9)
```

Un cappello sarà sempre creato con almeno una pallina. Gli argomenti passati all'oggetto cappello al momento della creazione dovrebbero essere convertiti in una variabile di istanza `contents`. `contents` dovrebbe essere una lista di stringhe contenente un elemento per ogni pallina nel cappello. Ogni elemento della lista dovrebbe essere un nome di colore che rappresenta una singola pallina di quel colore. Per esempio, se il tuo cappello è `{'red': 2, 'blue': 1}`, `contents` dovrebbe essere `['red', 'red', 'blue']`.

La classe `Hat` dovrebbe avere un metodo `draw` che accetta un argomento che indica il numero di palline da estrarre dal cappello. Questo metodo dovrebbe rimuovere palline a caso da `contents` e restituire quelle palline come una lista di stringhe. Le palline non dovrebbero tornare nel cappello durante l'estrazione, simile a un esperimento con urna senza reinserimento. Se il numero di palline da estrarre supera la quantità disponibile, restituisci tutte le palline.

Successivamente, crea una funzione `experiment` in `main.py` (non all'interno della classe `Hat`). Questa funzione dovrebbe accettare i seguenti argomenti:

- `hat`: un oggetto cappello contenente palline che dovrebbero essere copiate all'interno della funzione.
- `expected_balls`: un oggetto che indica il gruppo esatto di palline da tentare di estrarre dal cappello per l'esperimento. Per esempio, per determinare la probabilità di estrarre 2 palline blu e 1 pallina rossa dal cappello, imposta `expected_balls` su `{'blue':2, 'red':1}`.
- `num_balls_drawn`: il numero di palline da estrarre dal cappello in ogni esperimento.
- `num_experiments`: il numero di esperimenti da eseguire. (Più esperimenti esegui, più accurata sarà la probabilità approssimativa.)

La funzione `experiment` dovrebbe restituire una probabilità.

Per esempio, se vuoi determinare la probabilità di ottenere almeno due palline rosse e una pallina verde quando estrai cinque palline da un cappello contenente sei nere, quattro rosse e tre verdi. Per farlo, eseguirai `N` esperimenti, conterai quante volte `M` ottieni almeno due palline rosse e una pallina verde, e stimerai la probabilità come `M/N`. Ogni esperimento consiste nel partire da un cappello contenente le palline specificate, estrarre alcune palline e verificare se hai ottenuto le palline che stavi cercando di estrarre.

Ecco come chiameresti la funzione `experiment` basandoti sull'esempio sopra con 2000 esperimenti:

```py
hat = Hat(black=6, red=4, green=3)
probability = experiment(hat=hat,
                  expected_balls={'red':2,'green':1},
                  num_balls_drawn=5,
                  num_experiments=2000)
```

L'output sarebbe qualcosa del genere:

```bash
0.356
```

Poiché si basa su estrazioni casuali, la probabilità sarà leggermente diversa ogni volta che il codice viene eseguito.

_Suggerimento: considera di usare i moduli già importati in cima. Non inizializzare il seed di random all'interno del file._

Nota: apri la console del browser con F12 per vedere un output più dettagliato dei test.

# --hints--

La creazione dell'oggetto `hat` dovrebbe aggiungere il contenuto corretto.

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_class_contents(self):
        hat = probability_calculator.Hat(red=3,blue=2)
        actual = hat.contents
        expected = ["red","red","red","blue","blue"]
        self.assertEqual(actual, expected, 'Expected creation of hat object to add correct contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

Il metodo `draw` nella classe `hat` dovrebbe ridurre il numero di elementi nel contenuto.


```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)

class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_draw(self):
        hat = probability_calculator.Hat(red=5,blue=2)
        actual = hat.draw(2)
        expected = ['blue', 'red']
        self.assertEqual(actual, expected, 'Expected hat draw to return two random items from hat contents.')
        actual = len(hat.contents)
        expected = 5
        self.assertEqual(actual, expected, 'Expected hat draw to reduce number of items in contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

Il metodo `draw` dovrebbe comportarsi correttamente quando il numero di palline da estrarre è maggiore del numero di palline nel cappello.

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_draw_2(self):
        hat = probability_calculator.Hat(yellow=5,red=1,green=3,blue=9,test=1)
        actual = sorted(hat.draw(20))
        expected = sorted(['yellow', 'yellow', 'yellow', 'yellow', 'yellow', 'red', 'green', 'green', 'green', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'test'])
        self.assertEqual(actual, expected, 'Expected hat draw to return all items from hat contents.')
        actual = len(hat.contents)
        expected = 0
        self.assertEqual(actual, expected, 'Expected hat draw to leave no items in contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

Il metodo `experiment` dovrebbe restituire una probabilità diversa.


```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_prob_experiment(self):
        hat = probability_calculator.Hat(blue=3,red=2,green=6)
        probability = probability_calculator.experiment(hat=hat, expected_balls={"blue":2,"green":1}, num_balls_drawn=4, num_experiments=1000)
        actual = probability
        expected = 0.272
        self.assertAlmostEqual(actual, expected, delta = 0.01, msg = 'Expected experiment method to return a different probability.')
        hat = probability_calculator.Hat(yellow=5,red=1,green=3,blue=9,test=1)
        probability = probability_calculator.experiment(hat=hat, expected_balls={"yellow":2,"blue":3,"test":1}, num_balls_drawn=20, num_experiments=100)
        actual = probability
        expected = 1.0
        self.assertAlmostEqual(actual, expected, delta = 0.01, msg = 'Expected experiment method to return a different probability.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

# --seed--

## --seed-contents--

```py
import copy
import random

class Hat:
    pass

def experiment(hat, expected_balls, num_balls_drawn, num_experiments):
    pass
```

# --solutions--

```py
import copy
import random

class Hat:
    def __init__(self, **hat):
        self.hat = hat
        contents = []
        for i in hat:
            for j in range(hat[i]):
                contents.append(i)           
        self.contents = contents
               

    def draw(self, number): 
        drawn = []
        if number >= len(self.contents):
            drawn.extend(self.contents)
            self.contents = []
        else:
            for i in range(number):
                drawn.append(
                    self.contents.pop(random.randrange(len(self.contents)))
                )                
        return drawn

def experiment(hat, expected_balls, num_balls_drawn, num_experiments):    
    expected_balls_list = []
    drawn_list = []
    success = 0
    for i in expected_balls:
        for j in range(expected_balls[i]):
            expected_balls_list.append(i)
    for j in range(num_experiments):
        hat_copy = copy.deepcopy(hat)
        drawn_list.append(hat_copy.draw(num_balls_drawn))        
        exp_ball_list_copy = expected_balls_list[:]
        for k in range(len(drawn_list[j])):
            try:
                ind = exp_ball_list_copy.index(drawn_list[j][k])
                exp_ball_list_copy.pop(ind)
            except:
                continue
        if len(exp_ball_list_copy) == 0:
            success += 1        
    probability = success/num_experiments    
    return probability
```
