---
id: 66b59b12a745e10011158f7b
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

O próximo passo é criar mais algumas variáveis que serão usadas em mensagens futuras do bot.

Crie uma variável chamada `subject` e atribua a ela o valor de string `"JavaScript"`.

Em seguida, crie uma variável chamada `topic` e atribua a ela o valor de string `"strings"`.

# --hints--

Você deve ter uma variável chamada `subject`.

```js
assert.isNotNull(subject);
```

Você deve atribuir uma string à variável `subject`.

```js
assert.isString(subject);
```

Você deve atribuir a string `"JavaScript"` à variável `subject`.

```js
assert.equal(subject, "JavaScript");
```

Você deve ter uma variável chamada `topic`.

```js
assert.isNotNull(topic);
```

Você deve atribuir uma string à variável `topic`.

```js
assert.isString(topic);
```

Você deve atribuir a string `"strings"` à variável `topic`.

```js
assert.equal(topic, "strings");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

--fcc-editable-region--

--fcc-editable-region--
```
