---
id: 66b59be6ab830800c4df9146
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Agora, é hora de usar as variáveis que você criou no passo anterior.

Nos vídeos da aula, você aprendeu a trabalhar com template literals assim:

```js
const name = "John";
const age = 30;

// My name is John and I am 30 years old.
`My name is ${name} and I am ${age} years old.`;
```

Comece criando uma variável `sentence`.

Usando a sintaxe de template literals, atribua a string `Today, you will learn about [topic variable goes here] in [subject variable goes here].` à variável `sentence`.

Você substituirá os espaços reservados `[topic variable goes here]` e `[subject variable goes here]` pelas variáveis `topic` e `subject` usando a sintaxe `${}`.

Por fim, registre a variável `sentence` no console.

# --hints--

Você deve ter uma variável `sentence`.

```js
assert.isNotNull(sentence);
```

A variável `sentence` deve ser uma string.

```js
assert.isString(sentence);
```

Você deve usar a sintaxe de template literals para atribuir a string `Today, you will learn about [topic variable goes here] in [subject variable goes here].` à variável `sentence`. Consulte o exemplo se precisar de ajuda.

```js
assert.match(code, /Today,\s*you\s+will\s+learn\s+about\s+\$\{topic\}\s+in\s+\$\{subject\}\./);
```

Você deve registrar a variável `sentence` no console.

```js
assert.match(code,/console\.log\(\s*sentence\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

--fcc-editable-region--

--fcc-editable-region--
```
