---
id: 66b6efddeca35833cd6f0b03
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

Nos vídeos das aulas anteriores, você aprendeu a acessar caracteres em uma string assim:

```js
const firstName = "Jessica";
// returns "J"
firstName[0];
```

Lembre-se de que os números de índice começam em `0`. Portanto, a primeira letra em uma string estará sempre no índice `0`.

Comece adicionando outra instrução de `console`.

Dentro da instrução de `console`, exiba a primeira letra da variável `subject` usando notação de colchetes e o número de índice correto.

# --hints--

A instrução de `console` deve exibir a primeira letra da variável `subject` usando notação de colchetes e o índice `0`.

```js
assert.match(code, /console.log\(subject\[0\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
