---
id: 66b6d482bbb9e12f2e5ee1ae
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Para a próxima parte do projeto, o bot ensinará a trabalhar com a propriedade `length` da string.

Comece criando uma variável chamada `strLengthIntro`.

Em seguida, usando a sintaxe de template literals, atribua a string `Here is an example of using the length property on the word [subject].` à variável `strLengthIntro`.

Substitua `[subject]` pela variável `subject` como você fez anteriormente.

Finalmente, exiba o valor da variável `strLengthIntro` no console.

# --hints--

Você deve ter uma variável chamada `strLengthIntro`.

```js
assert.isNotNull(strLengthIntro);
```

A variável `strLengthIntro` deve ser uma string.

```js
assert.isString(strLengthIntro);
```

Você deve usar a sintaxe de template literals para atribuir a string `Here is an example of using the length property on the word [subject goes here].` à variável `strLengthIntro`.

```js
assert.match(code,/`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{subject\}\.`/);
```

Você deve exibir o valor da variável `strLengthIntro` no console.

```js
assert.match(code,/console\.log\(\s*strLengthIntro\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

--fcc-editable-region--

--fcc-editable-region--
```
