---
id: 587d8259367417b2b2512c85
title: реалізувати алгоритм сортування вибором
challengeType: 26
dashedName: implement-the-selection-sort-algorithm
---

# --description--

У цій лабораторній роботі ви реалізуєте алгоритм сортування вибором. Сортування вибором працює шляхом вибору мінімального значення у списку та обміну його з першим значенням у списку. Потім воно починає з другої позиції, вибирає найменше значення у залишку списку і обмінює його з другим елементом. Процес продовжується, ітеруючись по списку та обмінюючи елементи, доки не буде досягнуто кінця списку. Тепер список відсортований. Сортування вибором має квадратичну часову складність у всіх випадках.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні мати функцію `selectionSort`.
1. Функція `selectionSort` повинна приймати масив і повертати масив із тими ж елементами, але впорядкованими від найменшого до найбільшого.

# --hints--

`selectionSort` має бути функцією.

```js
assert.isFunction(selectionSort);
```

`selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` має повертати масив, який не змінений, крім порядку.

```js
assert.sameMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`selectionSort` має повертати відсортований масив (від найменшого до найбільшого).

```js
assert.sameOrderedMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
)
```

`selectionSort` не повинен використовувати вбудований метод `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
