---
id: 5e44412c903586ffb414c94c
title: 算術フォーマッターのプロジェクトを作成する
challengeType: 23
saveSubmissionToDB: true
forumTopicId: 462359
dashedName: build-an-arithmetic-formatter-project
---

# --description--

小学生は計算問題を縦に並べて解きやすくすることがよくあります。例えば、「235 + 52」は次のようになります：

```py
  235
+  52
-----
```

文字列のリストとして与えられる計算問題を受け取り、それらを縦に並べて横に並べて表示する`arithmetic_arranger`関数を完成させてください。関数はオプションで2番目の引数を受け取ることができます。2番目の引数が`True`に設定されている場合、答えも表示されます。

## 例

関数の呼び出し：

```py
arithmetic_arranger(["32 + 698", "3801 - 2", "45 + 43", "123 + 49"])
```

出力：

```py
   32      3801      45      123
+ 698    -    2    + 43    +  49
-----    ------    ----    -----
```

関数の呼び出し：

```py
arithmetic_arranger(["32 + 8", "1 - 3801", "9999 + 9999", "523 - 49"], True)
```

出力：

```py
  32         1      9999      523
+  8    - 3801    + 9999    -  49
----    ------    ------    -----
  40     -3800     19998      474
```

## ルール

与えられた問題が正しくフォーマットされていれば関数は正しい変換を返します。そうでなければ、ユーザーに意味のあるエラーを説明する**文字列**を**返します**。

- エラーを返す状況：
  - 関数に渡される問題が**多すぎる**場合。上限は**5つ**で、それ以上は次のエラーを返します：
    `'Error: Too many problems.'`
  - 関数が受け付ける演算子は**加算**と**減算**のみです。掛け算や割り算はエラーになります。この箇条書きにない他の演算子はテスト不要です。返されるエラーは：
    `"Error: Operator must be '+' or '-'."`
  - 各数字（オペランド）は数字のみで構成されている必要があります。そうでなければ、関数は次を返します：
    `'Error: Numbers must only contain digits.'`
  - 各オペランド（演算子の両側の数字）は最大4桁までです。そうでなければ、返されるエラー文字列は：
    `'Error: Numbers cannot be more than four digits.'`
- ユーザーが正しいフォーマットの問題を渡した場合、返す変換は以下のルールに従います：
  - 演算子と2つのオペランドのうち長い方の間には1つのスペースを入れます。演算子は2番目のオペランドと同じ行にあり、両方のオペランドは与えられた順序のまま（最初が上、2番目が下）です。
  - 数字は右揃えにします。
  - 各問題の間には4つのスペースを入れます。
  - 各問題の下にはダッシュを引きます。ダッシュは各問題の長さに合わせて引かれます。（上の例がどのようになるかを示しています）

注意：テストの詳細な出力を見るにはF12でブラウザのコンソールを開いてください。

# --hints--

`arithmetic_arranger(["3801 - 2", "123 + 49"])`は`  3801      123\n-    2    +  49\n------    -----`を返すべきです。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3801 - 2", "123 + 49"]), '  3801      123\\n-    2    +  49\\n------    -----')`);
  }
})
```

`arithmetic_arranger(["1 + 2", "1 - 9380"])`は`  1         1\n+ 2    - 9380\n---    ------`を返すべきです。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["1 + 2", "1 - 9380"]), '  1         1\\n+ 2    - 9380\\n---    ------')`);
  }
})
```

`arithmetic_arranger(["3 + 855", "3801 - 2", "45 + 43", "123 + 49"])`は`    3      3801      45      123\n+ 855    -    2    + 43    +  49\n-----    ------    ----    -----`を返すべきです。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 + 855", "3801 - 2", "45 + 43", "123 + 49"]), '    3      3801      45      123\\n+ 855    -    2    + 43    +  49\\n-----    ------    ----    -----')`);
  }
})
```

`arithmetic_arranger(["11 + 4", "3801 - 2999", "1 + 2", "123 + 49", "1 - 9380"])`は`  11      3801      1      123         1\n+  4    - 2999    + 2    +  49    - 9380\n----    ------    ---    -----    ------`を返すべきです。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["11 + 4", "3801 - 2999", "1 + 2", "123 + 49", "1 - 9380"]), '  11      3801      1      123         1\\n+  4    - 2999    + 2    +  49    - 9380\\n----    ------    ---    -----    ------')`);
  }
})
```

`arithmetic_arranger(["44 + 815", "909 - 2", "45 + 43", "123 + 49", "888 + 40", "653 + 87"])`は`'Error: Too many problems.'`を返すべきです。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["44 + 815", "909 - 2", "45 + 43", "123 + 49", "888 + 40", "653 + 87"]), 'Error: Too many problems.')`);
  }
})
```

`arithmetic_arranger(["3 / 855", "3801 - 2", "45 + 43", "123 + 49"])`は`"Error: Operator must be '+' or '-'."`を返すべきです。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 / 855", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Operator must be '+' or '-'.")`);
  }
})
```

`arithmetic_arranger(["24 + 85215", "3801 - 2", "45 + 43", "123 + 49"])`は`'Error: Numbers cannot be more than four digits.'`を返すべきです。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["24 + 85215", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Numbers cannot be more than four digits.")`);
  }
})
```

`arithmetic_arranger(["98 + 3g5", "3801 - 2", "45 + 43", "123 + 49"])`は`'Error: Numbers must only contain digits.'`を返すべきです。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["98 + 3g5", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Numbers must only contain digits.")`);
  }
})
```

`arithmetic_arranger(["3 + 855", "988 + 40"], True)`は`    3      988\n+ 855    +  40\n-----    -----\n  858     1028`を返すべきです。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 + 855", "988 + 40"], True), "    3      988\\n+ 855    +  40\\n-----    -----\\n  858     1028")`);
  }
})
```

`arithmetic_arranger(["32 - 698", "1 - 3801", "45 + 43", "123 + 49", "988 + 40"], True)`は`   32         1      45      123      988\n- 698    - 3801    + 43    +  49    +  40\n-----    ------    ----    -----    -----\n -666     -3800      88      172     1028`を返すべきです。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["32 - 698", "1 - 3801", "45 + 43", "123 + 49", "988 + 40"], True), "   32         1      45      123      988\\n- 698    - 3801    + 43    +  49    +  40\\n-----    ------    ----    -----    -----\\n -666     -3800      88      172     1028")`);
  }
})
```

# --seed--

## --seed-contents--

```py
def arithmetic_arranger(problems, show_answers=False):

    return problems

print(f'\n{arithmetic_arranger(["32 + 698", "3801 - 2", "45 + 43", "123 + 49"])}')
```

# --solutions--

```py
def arithmetic_arranger(problems, result=False):

	lin1 = ""
	lin2 = ""
	lin3 = ""
	lin4 = ""

	if len(problems) > 5:
		return 'Error: Too many problems.'
	for problem in problems:
		[num1, sym, num2] = problem.split()
		sign = ['+', '-']
		if sym not in sign:
			return ("Error: Operator must be '+' or '-'.")
		if len(num1) > 4 or len(num2) > 4:
			return ("Error: Numbers cannot be more than four digits.")
		if not num1.isnumeric() or not num2.isnumeric():
			return ("Error: Numbers must only contain digits.")

		lin1 += "  " + num1 + "    " if len(num1) >= len(
		    num2) else " " * (len(num2) + 2 - len(num1)) + num1 + "    "
		lin2 += sym + " " + num2 + "    " if len(num2) >= len(
		    num1) else sym + " " * (len(num1) - len(num2) + 1) + num2 + "    "
		nmax = (len(num1) + 2) if len(num1) >= len(num2) else (len(num2) + 2)
		lin3 += "-" * nmax + "    "
		op = int(num1) + int(num2) if sym == "+" else int(num1) - int(num2)
		lin4 += (" " * (nmax - len(str(op)))) + str(op) + "    "

	arranged_problems = lin1.rstrip() + "\n" + lin2.rstrip(
	) + "\n" + lin3.rstrip()
	if result:
		arranged_problems += "\n" + lin4.rstrip()

	return arranged_problems
```
