---
id: 673339d96c56d36e963e888e
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Neste workshop, você vai criar um catálogo de viveiro de plantas enquanto pratica os objetos Map e Set do JavaScript.

Para representar cada planta no seu catálogo, você usará objetos, cada um com três propriedades:

- `commonName`, o nome comum da planta
- `scientificName`, o nome científico (ou nome latino) da planta
- `cultivar`, a variedade da planta

Comece criando um objeto com as três propriedades acima para a `Spanish lavender` (nome científico `Lavandula stoechas`) de um cultivar chamado `Ballerina` e atribua-o a uma variável chamada `ballerina`.

# --hints--

Você deve ter uma variável chamada `ballerina`.

```js
assert.exists(ballerina);
```

Você deve atribuir um objeto à sua variável `ballerina`.

```js
assert.isObject(ballerina);
```

Seu objeto `ballerina` deve ter uma propriedade `commonName` definida como `"Spanish lavender"`.

```js
assert.strictEqual(ballerina.commonName, "Spanish lavender");
```

Seu objeto `ballerina` deve ter uma propriedade `scientificName` definida como `"Lavandula stoechas"`.

```js
assert.strictEqual(ballerina.scientificName, "Lavandula stoechas");
```

Seu objeto `ballerina` deve ter uma propriedade `cultivar` definida como `"Ballerina"`.

```js
assert.strictEqual(ballerina.cultivar, "Ballerina");
```

Seu objeto `ballerina` deve ter três chaves.

```js
assert.lengthOf(Object.keys(ballerina), 3);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
