---
id: 6733b5632fb854bc6c58dc0d
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Agora crie outro objeto, com as mesmas três propriedades, para representar outra cultivar de lavanda espanhola chamada `Pretty Polly` e atribua-a a uma variável chamada `prettyPolly`.

# --hints--

Você deve ter uma variável chamada `prettyPolly`.

```js
assert.exists(prettyPolly);
```

Você deve atribuir um objeto à sua variável `prettyPolly`.

```js
assert.isObject(prettyPolly);
```

Seu objeto `prettyPolly` deve ter uma propriedade `commonName` definida como `"Spanish lavender"`.

```js
assert.strictEqual(prettyPolly.commonName, "Spanish lavender");
```

Seu objeto `prettyPolly` deve ter uma propriedade `scientificName` definida como `"Lavandula stoechas"`.

```js
assert.strictEqual(prettyPolly.scientificName, "Lavandula stoechas");
```

Seu objeto `prettyPolly` deve ter uma propriedade `cultivar` definida como `"Pretty Polly"`.

```js
assert.strictEqual(prettyPolly.cultivar, "Pretty Polly");
```

Seu objeto `prettyPolly` deve ter três chaves.

```js
assert.lengthOf(Object.keys(prettyPolly), 3);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}
--fcc-editable-region--

--fcc-editable-region--
```
