---
id: 6733b5e6f210c1bd1afdf8b3
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Os objetos de planta restantes foram adicionados para você. Agora é hora de criar seu catálogo.

Como objetos JavaScript comuns, objetos `Map` contêm uma coleção de pares chave-valor. Para criar um objeto `Map`, você pode usar o construtor `Map`:

```js
const myMap = new Map();
```

Declare uma variável `catalog` e atribua a ela um objeto `Map`.

# --hints--

Você deve ter uma variável chamada `catalog`.

```js
assert.exists(catalog);
```

Você deve atribuir um objeto `Map` à sua variável `catalog`.

```js
assert.instanceOf(catalog, Map);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--

--fcc-editable-region--
```
