---
id: 6734a606e5968a5a19b77171
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

Continue adicionando as plantas restantes ao seu catálogo como fez nas etapas anteriores.

# --hints--

Você deve definir seu objeto `willowVale` como uma chave do seu mapa `catalog`.

```js
assert.isTrue(catalog.has(willowVale));
```

O valor da chave `willowVale` no mapa `catalog` deve ser um objeto com as propriedades `small`, `medium` e `large`, cada uma com o valor de um número inteiro positivo.

```js
const b = catalog.get(willowVale)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Você deve definir seu objeto `hidcote` como uma chave do seu mapa `catalog`.

```js
assert.isTrue(catalog.has(hidcote));
```

O valor da chave `hidcote` no mapa `catalog` deve ser um objeto com as propriedades `small`, `medium` e `large`, cada uma com o valor de um número inteiro positivo.

```js
const b = catalog.get(hidcote)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Você deve definir seu objeto `imperialGem` como uma chave do seu mapa `catalog`.

```js
assert.isTrue(catalog.has(imperialGem));
```

O valor da chave `imperialGem` no mapa `catalog` deve ser um objeto com as propriedades `small`, `medium` e `large`, cada uma com o valor de um número inteiro positivo.

```js
const b = catalog.get(imperialGem)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Você deve definir seu objeto `royalCrown` como uma chave do seu mapa `catalog`.

```js
assert.isTrue(catalog.has(royalCrown));
```

O valor da chave `royalCrown` no mapa `catalog` deve ser um objeto com as propriedades `small`, `medium` e `large`, cada uma com o valor de um número inteiro positivo.

```js
const b = catalog.get(royalCrown)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Você ainda deve ter a chave `ballerina` no seu mapa `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

Você ainda deve ter a chave `prettyPolly` no seu mapa `catalog`.

```js
assert.isTrue(catalog.has(prettyPolly));
```

Seu mapa `catalog` deve conter seis elementos.

```js
assert.lengthOf(catalog, 6);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
