---
id: 6734c10a327760a665b7d5b9
title: Passo 13
challengeType: 1
dashedName: step-13
---

# --description--

Nos próximos passos, você trabalhará em uma função que permitirá atualizar o catálogo quando você vender plantas.

Crie uma função chamada `sellPlants` que recebe três argumentos: um objeto planta, o tamanho do vaso e o número de vasos para vender, nesta ordem.

Quando o número especificado de potes for maior que os potes disponíveis, faça sua função retornar `Not enough <pot-size> size pots for <item-name>.`Apenas <pot-number> restante.`.Substitua `<pot-size>` pelo tamanho do vaso, `<pot-number>` pelos vasos restantes para esse tamanho e `<item-name>` pelo nome científico da planta seguido de um espaço e o nome do cultivar entre aspas simples.

Para testar se tudo funciona, registre o resultado de chamar sua função `sellPlants` com o objeto `ballerina`, `"small"` e `25` no console. 

# --hints--

Você não deve modificar seu mapa `catalog`.

```js
const testCatalog = new Map();
testCatalog.set(ballerina, { small: 20, medium: 15, large: 12 });
testCatalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
testCatalog.set(willowVale, { small: 3, medium: 5, large: 0 });
testCatalog.set(hidcote, { small: 33, medium: 13, large: 18 });
testCatalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
testCatalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
assert.deepEqual(catalog, testCatalog);
```

Você deve ter uma função chamada `sellPlants`.

```js
assert.isFunction(sellPlants);
```

Sua função `sellPlants` deve receber três argumentos.

```js
assert.lengthOf(sellPlants, 3);
```

`sellPlants(ballerina, "small", 25)` deve retornar `Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.`.

```js
assert.equal(sellPlants(ballerina, "small", 25), "Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.")

// prevent hardcoding
assert.equal(sellPlants(hidcote, "large", 95), "Not enough large size pots for Lavandula angustifolia 'Hidcote'. Only 18 left.")

assert.equal(sellPlants(royalCrown, "medium", 95), "Not enough medium size pots for Lavandula dentata 'Royal Crown'. Only 22 left.")
```

Você deve registrar `sellPlants(ballerina, "small", 25)` no console.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*sellPlants\s*\(\s*ballerina\s*,\s*("|')small\1\s*,\s*25\s*\)\s*\)/);
```

`sellPlants(ballerina, "small", 10)` deve retornar `undefined`.

```js
assert.isUndefined(sellPlants(ballerina, "small", 10));
// prevent hardcoding
assert.isUndefined(sellPlants(hidcote, "large", 5));
assert.isUndefined(sellPlants(royalCrown, "medium", 5));
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog)
```
