---
id: 6734eee1f9498cd90f9ae340
title: Passo 14
challengeType: 1
dashedName: step-14
---

# --description--

Quando houver vasos suficientes para vender, você precisa atualizar o catálogo subtraindo os vasos do tamanho especificado. Nesse caso, retorne `Catalog successfully updated.`

Modifique sua função para alcançar isso. Então, atualize sua chamada de função passando `10` como terceiro argumento em vez de `25`.

# --hints--

Você deve registrar `sellPlants(ballerina, "small", 10)` no console.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*sellPlants\s*\(\s*ballerina\s*,\s*("|')small\1\s*,\s*10\s*\)\s*\)/);
```

Você deve ter apenas uma chamada `sellPlants` no seu código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/(?<!function\s+)sellPlants\s*\(/g), 1)
```

`sellPlants(ballerina, "small", 25)` deve retornar `Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 10 left.`

```js
assert.equal(sellPlants(ballerina, "small", 25), "Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 10 left.")

// prevent hardcoding
assert.equal(sellPlants(hidcote, "large", 95), "Not enough large size pots for Lavandula angustifolia 'Hidcote'. Only 18 left.")
assert.equal(sellPlants(royalCrown, "medium", 95), "Not enough medium size pots for Lavandula dentata 'Royal Crown'. Only 22 left.")
```

`sellPlants(ballerina, "small", 10)` deve retornar `Catalog successfully updated.` e subtrair `10` da propriedade `small` do objeto associado à chave `ballerina` no catálogo.

```js
try {
    assert.equal(sellPlants(ballerina, "small", 10), "Catalog successfully updated.");
    assert.strictEqual(catalog.get(ballerina).small, 0);

    // prevent hardcoding
    assert.equal(sellPlants(hidcote, "large", 5), "Catalog successfully updated.");
    assert.strictEqual(catalog.get(hidcote).large, 13);

    assert.equal(sellPlants(royalCrown, "medium", 5), "Catalog successfully updated.");
    assert.strictEqual(catalog.get(royalCrown).medium, 17);
} finally {
    catalog.set(ballerina, { small: 10, medium: 15, large: 12 });
    catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
    catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
}
```

Você não deve modificar diretamente os pares chave-valor inseridos no seu mapa `catalog`.

```js
const testCatalog = new Map();
testCatalog.set(ballerina, { small: 10, medium: 15, large: 12 });
testCatalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
testCatalog.set(willowVale, { small: 3, medium: 5, large: 0 });
testCatalog.set(hidcote, { small: 33, medium: 13, large: 18 });
testCatalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
testCatalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
assert.deepEqual(catalog, testCatalog);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
--fcc-editable-region--
const sellPlants = (plant, size, potsNo) => {
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
}

console.log(sellPlants(ballerina, "small", 25));
--fcc-editable-region--
console.log(catalog)
```
