---
id: 6734a13a29acdf47f2b8e161
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Ao contrário dos objetos regulares, para os quais as chaves devem ser sempre strings, os objetos map permitem que você use qualquer tipo de dado como chaves.No seu catálogo, você vai usar pares de chave-valor objeto-objeto, associando cada um dos seus objetos planta com um objeto contendo o número de vasos disponíveis para diferentes tamanhos.

Modifique sua chamada `set` para usar seu objeto `ballerina` como chave.Para o valor, use um objeto contendo três propriedades, `small`, `medium` e `large`, cada uma com o valor de um número inteiro positivo. 

# --hints--

Você deve definir seu objeto `ballerina` como uma chave do seu mapa `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

O valor da chave `ballerina` no mapa `catalog` deve ser um objeto com as propriedades `small`, `medium` e `large`.

```js
const b = catalog.get(ballerina)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

A propriedade `small` do objeto associado à chave `ballerina` no mapa `catalog` deve ser um número inteiro positivo.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

A propriedade `medium` do objeto associado à chave `ballerina` no mapa `catalog` deve ser um número inteiro positivo.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

A propriedade `large` do objeto associado à chave `ballerina` no mapa `catalog` deve ser um número inteiro positivo.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Seu mapa `catalog` deve conter um único par chave-valor.

```js
assert.lengthOf(catalog, 1);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set("ballerina", ballerina);
--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
