---
id: 6734a2385c52c54d916a8b86
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Adicione outro par chave-valor ao seu catálogo.Use sua variável `prettyPolly` como a chave e faça o valor ser um objeto com as propriedades `small`, `medium` e `large` cujos valores são inteiros novamente.

# --hints--

Você deve definir seu objeto `prettyPolly` como uma chave do seu mapa `catalog`.

```js
assert.isTrue(catalog.has(prettyPolly));
```

O valor da chave `prettyPolly` no mapa `catalog` deve ser um objeto com as propriedades `small`, `medium` e `large`.

```js
const b = catalog.get(prettyPolly)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

A propriedade `small` do objeto associado à chave `prettyPolly` no mapa `catalog` deve ser um número inteiro positivo.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

A propriedade `medium` do objeto associado à chave `prettyPolly` no mapa `catalog` deve ser um número inteiro positivo.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

A propriedade `large` do objeto associado à chave `prettyPolly` no mapa `catalog` deve ser um número inteiro positivo.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Você ainda deve ter a chave `ballerina` no seu mapa `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

Seu mapa `catalog` deve conter dois elementos.

```js
assert.lengthOf(catalog, 2);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
