---
id: 67349ba230cff237c0f2bdf2
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Objetos Map lembram a ordem em que as chaves são inseridas no map.Para definir um par chave-valor no seu objeto map, você pode usar o método `set`.O seguinte define a chave `"name"` com o valor `"Jerry"` no objeto de mapa `myMap`.

```js
myMap.set("name", "Jerry");
```

Defina a chave `"ballerina"` com o valor do seu objeto `ballerina` no seu mapa `catalog`.Então, registre seu catálogo no console.

Observe que para ver seu objeto de mapa corretamente você terá que abrir o console do navegador.

# --hints--

Você deve definir a chave `"ballerina"` com o valor do seu objeto `ballerina` no seu mapa `catalog`.

```js
assert.deepEqual(catalog.get("ballerina"), ballerina)
```

Seu mapa `catalog` deve conter um único par chave-valor.

```js
assert.lengthOf(catalog, 1);
```

Você deve registrar seu mapa `catalog` no console.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*catalog\s*\)/)
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();

--fcc-editable-region--
```
