---
id: 0095d0567b6817114989a1dd
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستنفذ خوارزمية أقصر مسار في JavaScript. ستكتب دالة تحسب أقصر مسار بين العقد في رسم بياني موزون وتُرجع كل من المسافات والمسارات المتبعة.

على سبيل المثال، إذا كان هناك رسم بياني حيث ترتبط المدن بطرق ذات مسافات مختلفة، ستجد الخوارزمية أقصر مسار من مدينة إلى أخرى. إذا أردت السفر من المدينة A إلى المدينة D، قد تجد الخوارزمية أن المسار A ⇨ B ⇨ C ⇨ D (المجموع: 15 كم) أقصر من المسار المباشر A ⇨ D (20 كم).

للبدء، عرّف متغيرًا باسم `INF` وعيّن له القيمة `Infinity`. ستستخدمه لاحقًا للدلالة على عدم وجود اتصال مباشر بين عقدتين.

# --hints--

يجب أن تعرّف متغيرًا باسم `INF`.

```js
assert.isDefined(INF);
```

يجب أن يحتوي متغير `INF` على القيمة `Infinity`.

```js
assert.strictEqual(INF, Infinity);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
