---
id: 69c918841987bf16c1d1bac5
title: الخطوة 14
challengeType: 1
dashedName: step-14
---

# --description--

بعد فحص جميع العقد للعثور على أقرب عقدة غير مُزارة، يجب التعامل مع الحالة التي لا يتم فيها العثور على عقدة صالحة. يحدث هذا عندما تكون جميع العقد المتبقية غير قابلة للوصول.

بعد الحلقة الداخلية `for`، ولكن لا تزال داخل الحلقة الخارجية، أضف جملة `if` تتحقق مما إذا كان `current` يساوي تمامًا `-1`. إذا كان ذلك صحيحًا، استخدم `break` داخل كتلة الـ `if` لإيقاف الحلقة الخارجية مبكرًا.

# --hints--

يجب أن تضيف جملة `if` تتحقق مما إذا كان `current === -1`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)/
);
```

يجب أن تستخدم `break` داخل كتلة الـ `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)[\s\S]*?break/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
  
  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
