---
id: 69c918841987bf16c1d1baca
title: الخطوة 19
challengeType: 1
dashedName: step-19
---

# --description--

عندما تجد مسارًا أقصر إلى جار، يجب أيضًا تحديث المسار المسجل للوصول إليه.

في JavaScript، يمكنك إنشاء مصفوفة جديدة تجمع بين مصفوفة موجودة وعنصر جديد باستخدام معامل النشر:

```js
const newPath = [...existingArray, newElement];
```

داخل كتلة `if` المسكنة، حدّث `paths[nodeNo]` ليكون المسار إلى العقدة الحالية مع إضافة `nodeNo` في النهاية باستخدام صيغة النشر كما يلي: `[...paths[current], nodeNo]`.

# --hints--

يجب تحديث `paths[nodeNo]` داخل كتلة `if` المسكنة.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*\[\s*\w+\s*\]\s*=/
);
```

يجب تعيين `paths[nodeNo]` إلى `[...paths[current], nodeNo]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*\[\s*\w+\s*\]\s*=\s*\[\s*\.\.\.\s*paths\s*\[\s*current\s*\]\s*,\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
  
  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;
    
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          --fcc-editable-region--
          
          --fcc-editable-region--
        }
      }
    }
  }
}
```
