---
id: 69c918841987bf16c1d1bac4
title: الخطوة 13
challengeType: 1
dashedName: step-13
---

# --description--

إذا كان الشرط صحيحًا، فإن العقدة الحالية هي أفضل مرشح غير مُزار حتى الآن. يجب تحديث متغيرات التتبع لتعكس ذلك.

داخل جملة `if`، حدّث `minDistance` إلى `distances[nodeNo]` و`current` إلى `nodeNo`.

# --hints--

يجب تعيين `minDistance` إلى `distances[nodeNo]` داخل كتلة `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /minDistance\s*=\s*distances\s*\[\s*nodeNo\s*\]/
);
```

يجب تعيين `current` إلى `nodeNo` داخل كتلة `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /current\s*=\s*nodeNo/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
  
  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        --fcc-editable-region--
        
        --fcc-editable-region--
      }
    }
  }
}
```
