---
id: 69c918841987bf16c1d1bac1
title: الخطوة 10
challengeType: 1
dashedName: step-10
---

# --description--

الآن أضف الحلقة الرئيسية التي تُشغّل الخوارزمية. تعمل مرة واحدة لكل عقدة في الرسم البياني، وتختار أقرب عقدة غير مُزارة في كل مرة.

داخل الدالة `shortestPath`، أضف حلقة `for` تعمل `n` مرات باستخدام `i` كمتغير للحلقة.

في بداية جسم الحلقة، باستخدام `let`، أعلن عن متغيرين `minDistance` و`current` بقيم `INF` و`-1` على التوالي. ستتبع هذه المتغيرات أقرب عقدة غير مُزارة تم العثور عليها في كل تكرار.

# --hints--

يجب أن تحتوي دالة `shortestPath` على حلقة `for` تعمل `n` مرات. لا تنسَ إعلان `i` باستخدام `let` داخلها.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+i\s*=\s*0\s*;\s*i\s*<\s*n\s*;\s*i\+\+\s*\)/
);
```

يجب أن يكون لديك متغير `minDistance` بقيمة `INF` داخل الحلقة.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+minDistance\s*=\s*INF/
);
```

يجب أن يكون لديك متغير `current` بقيمة `-1` داخل الحلقة.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+current\s*=\s*-1/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
