---
id: 69c918841987bf16c1d1bac7
title: الخطوة 16
challengeType: 1
dashedName: step-16
---

# --description--

بما أن العقدة الحالية مُعلَمة بأنها مُزارة، يجب أن تنظر إلى جميع جيرانها لترى إذا كان بإمكانك الوصول إليهم بشكل أكثر كفاءة عبر العقدة الحالية.

بعد `visited[current] = true`، أضف حلقة `for` للجيران تُكرِّر عبر كل `nodeNo` من `0` إلى `n`. لاحظ أن هذه حلقة داخلية أخرى على نفس المستوى مثل الأولى.

داخل الحلقة، أعلن متغيرًا باسم `distance` وعيّنه إلى `matrix[current][nodeNo]`. هذا يعطيك وزن الحافة بين العقدة الحالية وجارها.

# --hints--

يجب أن تضيف حلقة `for` للجيران بعد تعليم العقدة الحالية بأنها مُزارة.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(/
);
```

يجب أن تُكرِّر حلقة `for` للجيران من `0` إلى `n`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(\s*(var|let)\s+\w+\s*=\s*0\s*;\s*\w+\s*<\s*n\s*;\s*\w+\+\+\s*\)/
);
```

يجب أن يكون لديك متغير `distance` مُعيَّن إلى `matrix[current][nodeNo]` داخل حلقة الجيران.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distance\s*=\s*matrix\s*\[\s*current\s*\]\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
  
  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;
    
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
