---
id: 69c937c73f23b1140bcd297a
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

معلمة `targetNode` اختيارية. عندما لا تُزوَّد، تُحسب الدالة أقصر المسارات من `startNode` إلى جميع العقد الأخرى في الرسم البياني.

أعطِ معلمة `targetNode` القيمة الافتراضية `null`.

# --hints--

يجب أن تكون معلمة `targetNode` في دالة `shortestPath` الخاصة بك بقيمة افتراضية `null`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /targetNode\s*=\s*null/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode) {
--fcc-editable-region--
  
}
```
