---
id: 69c918841987bf16c1d1baba
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

ستحتاج إلى مصفوفة ثنائية الأبعاد لتمثيل المصفوفة المجاورة للرسم البياني. كل قيمة تمثل الوزن (المسافة) للحافة بين عقدتين. القيمة `INF` تعني عدم وجود حافة مباشرة.

أنشئ متغيرًا آخر باسم `adjMatrix` وقم بتعيين مصفوفة ثنائية الأبعاد بالقيم التالية:

```js
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

يجب إنشاء متغير باسم `adjMatrix`.

```js
assert.isDefined(adjMatrix);
```

يجب أن تكون `adjMatrix` مصفوفة ثنائية الأبعاد تحتوي على 6 صفوف و6 أعمدة.

```js
assert.isArray(adjMatrix);
assert.lengthOf(adjMatrix, 6);
adjMatrix.forEach(row => assert.lengthOf(row, 6));
```

يجب أن تحتوي `adjMatrix` على القيم الصحيحة.

```js
assert.deepEqual(adjMatrix[0], [0, 5, 3, INF, 11, INF]);
assert.deepEqual(adjMatrix[1], [5, 0, 1, INF, INF, 2]);
assert.deepEqual(adjMatrix[2], [3, 1, 0, 1, 5, INF]);
assert.deepEqual(adjMatrix[3], [INF, INF, 1, 0, 9, 3]);
assert.deepEqual(adjMatrix[4], [11, INF, 5, 9, 0, INF]);
assert.deepEqual(adjMatrix[5], [INF, 2, INF, 3, INF, 0]);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
--fcc-editable-region--

--fcc-editable-region--
```
