---
id: 69c918841987bf16c1d1bace
title: الخطوة 23
challengeType: 1
dashedName: step-23
---

# --description--

بما أنك تعرف أن العقدة يمكن الوصول إليها وليست البداية، قم بتنسيق مسارها إلى سلسلة نصية قابلة للقراءة.

مصفوفة `paths[nodeNo]` تحتوي على متتالية مؤشرات العقد التي زُرت للوصول إلى `nodeNo`. يمكنك تحويلها إلى سلسلة نصية قابلة للقراءة مثل `0 -> 2 -> 3` باستخدام طريقة `.join()`.

داخل حلقة `for...of`، بعد جملة `if` التي تحتوي على `continue`، أعلن متغيرًا باسم `path` وعيّن له القيمة `paths[nodeNo].join(' -> ')`.

# --hints--

يجب أن يكون لديك متغير باسم `path` داخل حلقة `for...of`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /for\s*\(\s*(const|let)\s+nodeNo\s+of\s+targets\s*\)[\s\S]*?(const|let)\s+path\s*=/
);
```

يجب أن يحتوي متغير `path` على القيمة `paths[nodeNo].join(' -> ')`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /path\s*=\s*paths\s*\[\s*\w+\s*\]\s*\.\s*join\s*\(\s*['"]\s*->\s*['"]\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
  
  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;
    
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }
  
  const targets = targetNode !== null ? [targetNode] : [...Array(n).keys()];
  for (const nodeNo of targets) {
    if (nodeNo === startNode || distances[nodeNo] === INF) {
      continue;
    }
    --fcc-editable-region--
    
    --fcc-editable-region--
  }
}
```
