---
id: 69c918841987bf16c1d1babf
title: الخطوة 8
challengeType: 1
dashedName: step-8
---

# --description--

تحتاج أيضًا إلى تتبع المسار الفعلي الذي تم اتخاذه للوصول إلى كل عقدة. ستخزن قائمة من المسارات حيث يكون كل إدخال مصفوفة من فهارس العقد التي تمثل المسار المتبع.

في البداية، سيحتوي مسار كل عقدة على نفسها فقط. يمكنك إنشاء هذا الهيكل باستخدام `Array.from()`:

```js
Array.from({ length: n }, (_, i) => [i]);
```

هذا ينشئ مصفوفة من `n` عنصرًا، حيث يكون العنصر عند الفهرس `i` هو `[i]`.

لا تزال داخل دالة `shortestPath`، أنشئ متغيرًا باسم `paths` وعيّنه إلى `Array.from({ length: n }, (_, i) => [i])`.

# --hints--

يجب أن تنشئ متغيرًا باسم `paths` داخل دالة `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+paths\s*=/
);
```

يجب تهيئة متغير `paths` باستخدام `Array.from`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*=\s*Array\.from/
);
```

يجب أن يستخدم متغير `paths` `{ length: n }` كالمُعامل الأول لـ `Array.from`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\(\s*\{\s*length\s*:\s*n\s*\}/
);
```

المُعامل الثاني لـ `Array.from` يجب أن يكون استدعاء راجع `(_, i) => [i]`، حيث `_` هو عنصر نائب للعنصر غير المستخدم و`i` هو الفهرس. يجب أن يُرجع `[i]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\([\s\S]*?,\s*function\s*\(\s*\w+\s*,\s*(\w+)\s*\)\s*\{[\s\S]*?return\s*\[\s*\1\s*\]|Array\.from\s*\([\s\S]*?,\s*\(\s*\w+\s*,\s*\w+\s*\)\s*=>\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  --fcc-editable-region--
  
  --fcc-editable-region--
}
```
