---
id: 69c918841987bf16c1d1babc
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

تحتاج إلى تخزين عدد العقد في الرسم البياني.

داخل الدالة `shortestPath`، أنشئ متغيرًا باسم `n` وقم بتعيين القيمة `matrix.length` له.

# --hints--

يجب أن تنشئ متغيرًا `n` داخل الدالة `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+n\s*=/
);
```

يجب أن يحتوي متغير `n` على القيمة `matrix.length`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+n\s*=\s*matrix\.length/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
