---
id: 69c918841987bf16c1d1babd
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

يجب أن تتابع أقصر مسافة معروفة من عقدة البداية إلى كل عقدة أخرى. للبدء، افترض أن كل عقدة بعيدة إلى ما لا نهاية.

في JavaScript، يمكنك إنشاء مصفوفة مملوءة مسبقًا بقيمة باستخدام `new Array(n).fill(value)`. على سبيل المثال، `new Array(3).fill(0)` تُنشئ `[0, 0, 0]`.

داخل الدالة `shortestPath`، أنشئ متغيرًا باسم `distances` وعيّنه إلى `new Array(n).fill(INF)`.

# --hints--

يجب أن تنشئ متغيرًا باسم `distances` داخل `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distances\s*=/
);
```

يجب تعيين متغير `distances` إلى `new Array(n).fill(INF)`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*INF\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
