---
id: 612e6afc009b450a437940a1
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Begin with the basic HTML structure. Add a `DOCTYPE` declaration and `html`, `head`, `body`, and `title` elements.

اجعل الإنجليزية لغة هذه الصحفة. اجعل العنوان هو `Piano` باستعمال الوسم `title`.

# --hints--

يجب أن يحتوي الكود الخاص بك على مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

عليك أن تضيف مسافة بعد مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

عليك أن تعرّف أن نوع مستندك هو `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

ينبغي إغلاق إعلان `DOCTYPE` بواسطة <code dir="ltr">&gt;</code> بعد تعيين النوع.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

عنصر `html` الخاص بك يجب أن يكون له opening tag معه سمة `lang` تساوي `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصرك `html` على علامة أغلاق (closing tag).

```js
assert(code.match(/<\/html\s*>/));
```

تعريف `DOCTYPE` الخاص بك يجب أن يكون عند بداية HTML الخاص بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

يجب أن يكون لديك الـ opening tag للـ `head`.

```js
assert(code.match(/<head\s*>/i));
```

يجب أن يكون لديك الـ closing tag للـ `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

يجب أن يكون لديك الـ opening tag للـ `body`.

```js
assert(code.match(/<body\s*>/i));
```

يجب أن يكون لديك الـ closing tag للـ `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

يجب أن يكون عناصر `head` و `body` متلازمين (siblings).

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

عنصر `head` يجب أن يكون داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

عنصر `body` يجب أن يكون داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

الكود يجب أن يكون فيه عنصر `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

مشروعك يجب أن يكون له العنوان `Piano`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'piano')
```

تذكر، الإملاء وتصغير وتكبير الحروف (casing) مهمان في العنوان.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Piano');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
