---
id: 612ea4c4993aba52ab4aa32e
title: الخطوة 18
challengeType: 0
dashedName: step-18
---

# --description--

Now it is time to use the pseudo-selectors you prepared for earlier. To create the black keys, add a new `.key.black--key::after` selector. This will target the elements with the class `key black--key`, and select the pseudo-element after these elements in the HTML.

في المنتقي الجديد، عيّن `background-color` إلى `#1d1e22`. قم أيضا بتعيين خاصية `content` إلى `""`. وهذا سيجعل الـ pseudo-elements فارغة.

يتم استخدام خاصية `content` لتعيين أو تجاوز (override) محتوى العنصر. بشكل افتراضي، الـ psuedo-elements التي تم إنشاؤها بواسطة `::before` و `::after` هي فارغة، ولن تعرض (render) هذه العناصر علي الصفحة. سيؤدي تعيين خاصية `content` إلى سلسلة فارغة `""` إلى ضمان عرض العنصر علي الصفحة بينما يظل فارغاً.

إذا كنت ترغب في التجربة، حاول إزالة خاصية `background-color` ووضع قيم مختلفة لخاصية `content`، مثل `"♥"`. تذكر أن تتراجع (undo) عن هذه التغييرات عندما يتم إجراء الاختبارات.

# --hints--

يجب أن يكون لديك منتقي `.key.black--key::after`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

يجب أن يحتوي منتقي `.key.black--key::after` على خاصية `background-color` بقيمة `#1d1e22`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor === 'rgb(29, 30, 34)');
```

يجب أن يحتوي منتقي `.key.black--key::after` على خاصية `content` بقيمة `""`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content === '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
