---
id: 68dbc6b7e8f794a2ae69bc74
title: Tekeleza algoriti ya Malkia N
challengeType: 27
dashedName: implement-the-n-queens-algorithm
---

# --description--

Changamoto ya Malkia N inakuomba uweke malkia N kwenye ubao wa chess wa N×N kwa njia ambayo malkia wawili hawashambuliani (hakuna wawili wanaoshiriki safu, nguzo, au diagonal).

Kwa mfano, ikiwa kuna ubao wa 4x4, mpangilio halali mmoja ni:

```md
[1, 3, 0, 2]
```

Hii inamaanisha kwamba katika safu ya 0, malkia amewekwa katika nguzo ya 1; katika safu ya 1, malkia amewekwa katika nguzo ya 3; katika safu ya 2, malkia amewekwa katika nguzo ya 0; na katika safu ya 3, malkia amewekwa katika nguzo ya 2.

Kwa mtazamo wa macho, mpangilio huu unaonekana kama:

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

Ambapo `Q` inawakilisha malkia na `.` inawakilisha mraba tupu.

Katika maabara hii, utatekeleza suluhisho la changamoto ya Malkia N kwa kutumia mbinu ya utafutaji wa kina kwanza.


**Lengo**: Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kitendakazi kinachoitwa `dfs_n_queens`.
2. Kitendakazi kinapaswa kupokea hoja moja tu: nambari kamili `n`.
3. Ikiwa `n` ni ndogo kuliko `1`, kitendakazi kinapaswa kurudisha orodha tupu (`[]`).
4. Kitendakazi kinapaswa kurudisha orodha ya suluhisho; kila suluhisho ni orodha yenye urefu wa `n`, ambapo kipengele kilicho kwenye kielezo `i` ni kielezo cha nguzo (kuanzia 0) cha malkia katika safu `i`.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `dfs_n_queens` kinachopokea hoja moja.

```js
({ test: () => runPython(`
  import inspect
  assert inspect.isfunction(dfs_n_queens)
  sig = inspect.signature(dfs_n_queens)
  assert len(sig.parameters) == 1
`) })
```

Ikiwa `n` ni ndogo kuliko `1`, kitendakazi kinapaswa kurudisha orodha tupu.

```js
({ test: () => runPython(`
  assert dfs_n_queens(0) == []
  assert dfs_n_queens(-1) == []
  assert dfs_n_queens(-5) == []
`) })
```

Kitendakazi kinapaswa kurudisha orodha ya suluhisho, ambapo kila suluhisho ni orodha yenye urefu wa `n`.

```js
({ test: () => runPython(`
  result = dfs_n_queens(4)
  assert isinstance(result, list)
  for solution in result:
    assert isinstance(solution, list)
    assert len(solution) == 4
`) })
```

`dfs_n_queens(1)` inapaswa kurudisha `[[0]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(1) == [[0]]
`) })
```

`dfs_n_queens(2)` inapaswa kurudisha `[]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(2) == []
`) })
```

`dfs_n_queens(3)` inapaswa kurudisha `[]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(3) == []
`) })
```

`dfs_n_queens(4)` inapaswa kurudisha `[[1, 3, 0, 2], [2, 0, 3, 1]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(4) == [[1, 3, 0, 2], [2, 0, 3, 1]]
`) })
```

`dfs_n_queens(5)` inapaswa kurudisha `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(5) == [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]
`) })
```

`len(dfs_n_queens(5))` inapaswa kuwa `10`.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(5)) == 10
`) })
```

`len(dfs_n_queens(8))` inapaswa kuwa `92`.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(8)) == 92
`) })
```

`dfs_n_queens` inapaswa kurudisha matokeo sahihi.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(7)) == 40
  assert dfs_n_queens(6) == [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def dfs_n_queens(n):
    if n < 1:
        return []

    results = []

    def is_safe(queens, row, col):
        for r, c in enumerate(queens):
            if c == col or abs(row - r) == abs(col - c):
                return False
        return True

    def dfs(queens):
        row = len(queens)
        if row == n:
            results.append(queens[:])
            return
        for col in range(n):
            if is_safe(queens, row, col):
                queens.append(col)
                dfs(queens)
                queens.pop()

    dfs([])
    return results
```
