---
id: 56533eb9ac21ba0edf2244b3
title: 创建一个摄氏度转华氏度的转换器
challengeType: 26
dashedName: build-a-celsius-to-fahrenheit-converter
---

# --description--

在本实验中，你将编写一个将摄氏温度转换为华氏温度的函数。 将摄氏度换算成华氏度的公式是：

```js
fahrenheit = celsius * (9/5) + 32
```

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应创建一个名为 `convertCtoF` 的函数。
2. `convertCtoF` 应该接受一个单一的数字参数，该参数是摄氏温度。
3. `convertCtoF` 应返回一个数字。

# --hints--

你应创建一个名为 `convertCtoF` 的函数。

```js
assert.isFunction(convertCtoF);
```

`convertCtoF` 应只接受一个参数。

```js
assert.lengthOf(convertCtoF, 1);
```

`convertCtoF(0)` 应该返回一个数字。

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` 应该返回一个值 `-22`。

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` 应该返回一个值 `14`。

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` 应返回一个值 `32`。

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` 应返回值 `68`。

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` 应返回一个值为 `86`。

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
