---
id: 5e44412c903586ffb414c94c
title: Створіть проєкт арифметичного форматування
challengeType: 23
saveSubmissionToDB: true
forumTopicId: 462359
dashedName: build-an-arithmetic-formatter-project
---

# --description--

Учні початкової школи часто розташовують арифметичні задачі вертикально, щоб їх було легше розв’язувати. Наприклад, "235 + 52" стає:

```py
  235
+  52
-----
```

Завершіть функцію `arithmetic_arranger`, яка приймає список рядків із арифметичними задачами і повертає задачі, розташовані вертикально та поруч одна з одною. Функція має необов’язково приймати другий аргумент. Якщо другий аргумент встановлено в `True`, потрібно відобразити відповіді.

## Приклад

Виклик функції:

```py
arithmetic_arranger(["32 + 698", "3801 - 2", "45 + 43", "123 + 49"])
```

Вихідні дані:

```py
   32      3801      45      123
+ 698    -    2    + 43    +  49
-----    ------    ----    -----
```

Виклик функції:

```py
arithmetic_arranger(["32 + 8", "1 - 3801", "9999 + 9999", "523 - 49"], True)
```

Вихідні дані:

```py
  32         1      9999      523
+  8    - 3801    + 9999    -  49
----    ------    ------    -----
  40     -3800     19998      474
```

## Правила

Функція поверне правильне форматування, якщо надані задачі оформлені правильно, інакше вона **поверне** **рядок**, який описує помилку, зрозумілу користувачу.

- Ситуації, які призведуть до помилки:
  - Якщо функції передано **занадто багато задач**. Ліміт — **п’ять**, більше за це поверне:
    `'Error: Too many problems.'`
  - Допустимі оператори — **додавання** та **віднімання**. Множення та ділення викличуть помилку. Інші оператори перевіряти не потрібно. Помилка буде:
    `"Error: Operator must be '+' or '-'."`
  - Кожне число (операнд) має містити лише цифри. Інакше функція поверне:
    `'Error: Numbers must only contain digits.'`
  - Кожен операнд (число з будь-якого боку оператора) має максимум чотири цифри. Інакше повернеться рядок з помилкою:
    `'Error: Numbers cannot be more than four digits.'`
- Якщо користувач надав задачі у правильному форматі, повернене форматування відповідатиме таким правилам:
  - Між оператором і найдовшим із двох операндів має бути один пробіл, оператор розташовується на тій же лінії, що й другий операнд, обидва операнди йдуть у тому ж порядку, що й у вхідних даних (перший зверху, другий знизу).
  - Числа мають бути вирівняні по правому краю.
  - Між кожною задачею має бути чотири пробіли.
  - Під кожною задачею мають бути тире. Тире мають тягнутися на всю довжину кожної задачі окремо. (Приклад вище показує, як це має виглядати.)

Примітка: відкрийте консоль браузера за допомогою F12, щоб побачити більш докладний вивід тестів.

# --hints--

`arithmetic_arranger(["3801 - 2", "123 + 49"])` має повертати `  3801      123\n-    2    +  49\n------    -----`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3801 - 2", "123 + 49"]), '  3801      123\\n-    2    +  49\\n------    -----')`);
  }
})
```

`arithmetic_arranger(["1 + 2", "1 - 9380"])` має повертати `  1         1\n+ 2    - 9380\n---    ------`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["1 + 2", "1 - 9380"]), '  1         1\\n+ 2    - 9380\\n---    ------')`);
  }
})
```

`arithmetic_arranger(["3 + 855", "3801 - 2", "45 + 43", "123 + 49"])` має повертати `    3      3801      45      123\n+ 855    -    2    + 43    +  49\n-----    ------    ----    -----`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 + 855", "3801 - 2", "45 + 43", "123 + 49"]), '    3      3801      45      123\\n+ 855    -    2    + 43    +  49\\n-----    ------    ----    -----')`);
  }
})
```

`arithmetic_arranger(["11 + 4", "3801 - 2999", "1 + 2", "123 + 49", "1 - 9380"])` має повертати `  11      3801      1      123         1\n+  4    - 2999    + 2    +  49    - 9380\n----    ------    ---    -----    ------`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["11 + 4", "3801 - 2999", "1 + 2", "123 + 49", "1 - 9380"]), '  11      3801      1      123         1\\n+  4    - 2999    + 2    +  49    - 9380\\n----    ------    ---    -----    ------')`);
  }
})
```

`arithmetic_arranger(["44 + 815", "909 - 2", "45 + 43", "123 + 49", "888 + 40", "653 + 87"])` має повертати `'Error: Too many problems.'`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["44 + 815", "909 - 2", "45 + 43", "123 + 49", "888 + 40", "653 + 87"]), 'Error: Too many problems.')`);
  }
})
```

`arithmetic_arranger(["3 / 855", "3801 - 2", "45 + 43", "123 + 49"])` має повертати `"Error: Operator must be '+' or '-'."`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 / 855", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Operator must be '+' or '-'.")`);
  }
})
```

`arithmetic_arranger(["24 + 85215", "3801 - 2", "45 + 43", "123 + 49"])` має повертати `'Error: Numbers cannot be more than four digits.'`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["24 + 85215", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Numbers cannot be more than four digits.")`);
  }
})
```

`arithmetic_arranger(["98 + 3g5", "3801 - 2", "45 + 43", "123 + 49"])` має повертати `'Error: Numbers must only contain digits.'`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["98 + 3g5", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Numbers must only contain digits.")`);
  }
})
```

`arithmetic_arranger(["3 + 855", "988 + 40"], True)` має повертати `    3      988\n+ 855    +  40\n-----    -----\n  858     1028`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 + 855", "988 + 40"], True), "    3      988\\n+ 855    +  40\\n-----    -----\\n  858     1028")`);
  }
})
```

`arithmetic_arranger(["32 - 698", "1 - 3801", "45 + 43", "123 + 49", "988 + 40"], True)` має повертати `   32         1      45      123      988\n- 698    - 3801    + 43    +  49    +  40\n-----    ------    ----    -----    -----\n -666     -3800      88      172     1028`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["32 - 698", "1 - 3801", "45 + 43", "123 + 49", "988 + 40"], True), "   32         1      45      123      988\\n- 698    - 3801    + 43    +  49    +  40\\n-----    ------    ----    -----    -----\\n -666     -3800      88      172     1028")`);
  }
})
```

# --seed--

## --seed-contents--

```py
def arithmetic_arranger(problems, show_answers=False):

    return problems

print(f'\n{arithmetic_arranger(["32 + 698", "3801 - 2", "45 + 43", "123 + 49"])}')
```

# --solutions--

```py
def arithmetic_arranger(problems, result=False):

	lin1 = ""
	lin2 = ""
	lin3 = ""
	lin4 = ""

	if len(problems) > 5:
		return 'Error: Too many problems.'
	for problem in problems:
		[num1, sym, num2] = problem.split()
		sign = ['+', '-']
		if sym not in sign:
			return ("Error: Operator must be '+' or '-'.")
		if len(num1) > 4 or len(num2) > 4:
			return ("Error: Numbers cannot be more than four digits.")
		if not num1.isnumeric() or not num2.isnumeric():
			return ("Error: Numbers must only contain digits.")

		lin1 += "  " + num1 + "    " if len(num1) >= len(
		    num2) else " " * (len(num2) + 2 - len(num1)) + num1 + "    "
		lin2 += sym + " " + num2 + "    " if len(num2) >= len(
		    num1) else sym + " " * (len(num1) - len(num2) + 1) + num2 + "    "
		nmax = (len(num1) + 2) if len(num1) >= len(num2) else (len(num2) + 2)
		lin3 += "-" * nmax + "    "
		op = int(num1) + int(num2) if sym == "+" else int(num1) - int(num2)
		lin4 += (" " * (nmax - len(str(op)))) + str(op) + "    "

	arranged_problems = lin1.rstrip() + "\n" + lin2.rstrip(
	) + "\n" + lin3.rstrip()
	if result:
		arranged_problems += "\n" + lin4.rstrip()

	return arranged_problems
```
