---
id: 68e4276483b0c71658aab1ce
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

문자열의 시작 부분에 있는 공백만 제거하고 싶을 때가 있습니다. 이럴 때 `trimStart()` 메서드가 유용합니다.

예를 들어 보겠습니다:

```js
const sentence = "  I love learning.  ";
sentence.trimStart(); // "I love learning.  "
```

문자열은 불변임을 기억하세요. 이 메서드는 결과를 가진 새 문자열을 만듭니다. 원래 문자열을 수정하지 않습니다.

변수 `trimmedStart`을 만들고 `trimStart()` 변수에 `userInput` 메서드를 사용한 결과를 할당하세요.

# --hints--

`trimmedStart` 변수가 있어야 합니다.

```js
assert.exists(trimmedStart);
```

`trimmedStart` 변수는 문자열이어야 합니다.

```js
assert.isString(trimmedStart);
```

`userInput.trimStart()`을 `trimmedStart` 변수에 할당해야 합니다.

```js
assert.equal(trimmedStart, userInput.trimStart());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);
--fcc-editable-region--

--fcc-editable-region--
```
