---
id: 68e4276483b0c71658aab1cc
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

`userInput` 값의 문자열 양쪽 끝에 여분의 공백이 있습니다. 그 공백을 제거한 새 문자열을 만들려면 `trim()` 메서드를 사용할 수 있습니다.  

예시는 다음과 같습니다:

```js
const sentence = "  Learning about the trim method.  ";
sentence.trim(); // "Learning about the trim method."
```

`cleanedInput`라는 변수를 만들고 `trim()` 변수에 `userInput` 메서드를 사용한 결과를 할당하세요.

# --hints--

`cleanedInput` 변수가 있어야 합니다.

```js
assert.exists(cleanedInput);
```

`cleanedInput` 변수는 문자열이어야 합니다.

```js
assert.isString(cleanedInput);
```

`userInput.trim()`을 `cleanedInput` 변수에 할당해야 합니다.

```js
assert.equal(cleanedInput, userInput.trim());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);
--fcc-editable-region--

--fcc-editable-region--
```
