---
id: 68e4276483b0c71658aab1d3
title: 14단계
challengeType: 1
dashedName: step-14
---

# --description--

이전 수업과 워크숍에서 기억하셨겠지만, `slice()` 메서드를 사용해 문자열의 일부를 추출할 수 있습니다.

다시 한 번 알려드립니다:

```js
const str = "abcd";
str.slice(0, 3); // "abc"
```

빈 문자열을 `camelCasedVersion`로 바꾸어 기존 `lowercaseWord.slice(0, 5)` 할당을 변경하세요.

이제 콘솔에서 `"camel"`라는 단어가 보여야 합니다.

# --hints--

`camelCasedVersion` 변수는 빈 문자열이면 안 됩니다.

```js
assert.isNotEmpty(camelCasedVersion);
```

`lowercaseWord.slice(0, 5)` 변수에 `camelCasedVersion`를 할당해야 합니다.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5));
```

문자열 리터럴 `"camel"`를 `camelCasedVersion`에 할당하면 안 됩니다. 올바른 방법을 위해 지침을 다시 읽어보세요.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camel(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = ""
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
