---
id: 68e4276483b0c71658aab1d1
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

모든 글자가 소문자인 새 문자열을 만들어야 한다면 `toLowerCase()` 메서드를 사용할 수 있습니다.  

먼저 `lowerCaseInput`이라는 변수를 만들고 `toLowerCase()` 변수에 `cleanedInput` 메서드를 사용한 결과를 할당하세요. 

그다음 문자열 `"Result of using the toLowerCase() method:"`를 콘솔에 출력하세요.

마지막으로 `console.log()` 변수에 대해 또 다른 `lowerCaseInput`를 추가하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`lowerCaseInput` 변수가 있어야 합니다.

```js
assert.exists(lowerCaseInput);
```

`lowerCaseInput` 변수는 문자열이어야 합니다.

```js
assert.isString(lowerCaseInput);
```

`cleanedInput.toLowerCase()`를 `lowerCaseInput` 변수에 할당해야 합니다.

```js
assert.equal(lowerCaseInput, cleanedInput.toLowerCase());
```

문자열 `"Result of using the toLowerCase() method:"`를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[10], "Result of using the toLowerCase() method:");
```

`lowerCaseInput` 변수를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[11], lowerCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);


const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
