---
id: 68e4276483b0c71658aab1d0
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

다른 대소문자 조합으로 새 문자열을 만들어야 할 때가 있습니다.

모든 문자를 대문자로 변환한 새 문자열을 만들려면 `toUpperCase()` 메서드를 사용할 수 있습니다.

예를 들어 보겠습니다:

```js
const sentence = "I love to code!";
sentence.toUpperCase(); // "I LOVE TO CODE!"
```

`upperCaseInput`이라는 변수를 만들고 `toUpperCase()` 변수에 `cleanedInput` 메서드를 사용한 결과를 할당하세요.

# --hints--

`upperCaseInput` 변수를 만들어야 합니다.

```js
assert.exists(upperCaseInput);
```

`upperCaseInput` 변수는 문자열이어야 합니다.

```js
assert.isString(upperCaseInput);
```

`cleanedInput.toUpperCase()`을 `upperCaseInput` 변수에 할당해야 합니다.

```js
assert.equal(upperCaseInput, cleanedInput.toUpperCase());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

--fcc-editable-region--

--fcc-editable-region--
```
