---
id: 68e4276483b0c71658aab1d2
title: 13단계
challengeType: 1
dashedName: step-13
---

# --description--

새로운 문자열을 만들 때 여러 메서드를 사용해야 하는 경우가 있습니다. 다음 몇 단계에서는 지금까지 배운 메서드 중 일부를 사용해 카멜 케이스 명명 관례를 따르는 새로운 문자열을 만들 것입니다.

먼저 `lowercaseWord`라는 변수를 만들고 문자열 `"camelcase"`을 할당하세요.

그 아래 `lowercaseWord` 변수 아래에 `camelCasedVersion`라는 또 다른 변수를 만들고 지금은 빈 문자열을 할당하세요.

그 아래 `camelCasedVersion` 변수 아래에 `console.log()`를 문자열 `"Camel cased version:"`과 함께 추가하고, `console.log()` 변수에 대해 또 다른 `camelCasedVersion`을 추가하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`lowercaseWord`라는 변수가 있어야 합니다.

```js
assert.exists(lowercaseWord);
```

`lowercaseWord`는 문자열이어야 합니다.

```js
assert.isString(lowercaseWord);
```

`"camelcase"` 변수에 문자열 `lowercaseWord`을 할당해야 합니다.

```js
assert.equal(lowercaseWord, "camelcase");
```

`camelCasedVersion` 변수가 있어야 합니다.

```js
assert.exists(camelCasedVersion);
```

`camelCasedVersion` 변수는 문자열이어야 합니다.

```js
assert.isString(camelCasedVersion);
```

`camelCasedVersion` 변수에 빈 문자열을 할당해야 합니다.

```js
assert.isEmpty(camelCasedVersion);
```

문자열 `"Camel cased version:"`을 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[12], "Camel cased version:");
```

`camelCasedVersion` 변수를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[13], camelCasedVersion);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
