---
id: 68e4276483b0c71658aab1ca
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 문자열을 포맷하는 데 사용되는 다양한 문자열 메서드를 다루는 연습을 합니다.

먼저 `userInput`라는 변수를 만들고 다음 값을 할당하세요: `"   Hello World!   "`.

# --hints--

`userInput` 변수가 있어야 합니다.

```js
assert.exists(userInput);
```

`userInput` 변수는 문자열이어야 합니다.

```js
assert.isString(userInput);
```

`"   Hello World!   "`를 `userInput` 변수에 할당하세요. `H` 앞에 공백 3개, `!` 뒤에 공백 3개가 있는지 확인하세요.

```js
assert.equal(userInput, "   Hello World!   ");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
