---
id: 68e4276483b0c71658aab1cd
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

이제 결과를 확인할 시간입니다.

먼저 `console.log()`에 문자열 `"Result of trimming whitespace from both ends:"`을 추가하세요.

그다음 그 아래 `console.log()`에 `console.log()` 변수를 가진 또 다른 `cleanedInput`을 추가하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

문자열 `"Result of trimming whitespace from both ends:"`를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[2], "Result of trimming whitespace from both ends:");
```

`cleanedInput` 변수를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[3], cleanedInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
--fcc-editable-region--

--fcc-editable-region--
```
