---
id: 68ed6079fe0e3e98b94a1b75
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

다음으로, 문자열 `console.log()`을 가진 `"Original input:"`을 추가하세요. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

문자열 `"Original input:"`을 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[0], "Original input:");
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
--fcc-editable-region--

--fcc-editable-region--
```
