---
id: 68f1179eaea3bb59ba434386
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

문자열 끝의 공백을 제거하려면 `trimEnd()` 메서드를 사용할 수 있습니다.

예시는 다음과 같습니다:

```js
const sentence = "  I love learning.  ";
sentence.trimEnd(); // "  I love learning."
```

`trimmedEnd`이라는 변수를 만들고 `trimEnd()` 변수에 `userInput` 메서드를 사용한 결과를 할당하세요.

# --hints--

`trimmedEnd` 변수가 있어야 합니다.

```js
assert.exists(trimmedEnd);
```

`trimmedEnd` 변수는 문자열이어야 합니다.

```js
assert.isString(trimmedEnd);
```

`userInput.trimEnd()`을 `trimmedEnd` 변수에 할당해야 합니다.

```js
assert.equal(trimmedEnd, userInput.trimEnd());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

--fcc-editable-region--

--fcc-editable-region--
```
