---
id: 68f3f3d65d51d80ac0b412fc
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

이제 문자열 시작 부분의 공백을 제거한 결과를 확인할 차례입니다.

`trimmedStart` 변수 아래에 `console.log()` 문자열 `"After using the trimStart() method, leading spaces removed:"`를 추가하세요.

그다음 `console.log()` 아래에 `console.log()` 변수를 가진 또 다른 `trimmedStart`를 추가하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

문자열 `"After using the trimStart() method, leading spaces removed:"`를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[4], "After using the trimStart() method, leading spaces removed:");
```

`trimmedStart` 변수를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[5], trimmedStart);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
--fcc-editable-region--

--fcc-editable-region--
```
