---
id: 68f3f890010c870e17bbe754
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

이제 결과를 확인할 시간입니다.

`upperCaseInput` 변수 아래에 문자열 `console.log()`를 가진 `"Result of using the toUpperCase() method:"`을 추가하세요.

그다음 `console.log()` 아래에 `console.log()` 변수를 가진 또 다른 `upperCaseInput`를 추가하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

문자열 `"Result of using the toUpperCase() method:"`를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[8], "Result of using the toUpperCase() method:");
```

`upperCaseInput` 변수를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[9], upperCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
--fcc-editable-region--

--fcc-editable-region--
```
