---
id: 68f3f5ca1d2a490c97f14f03
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

이제 문자열 끝의 공백을 제거한 결과를 확인할 시간입니다.

`trimmedEnd` 변수 아래에 `console.log()` 문자열 `"After using the trimEnd() method, trailing spaces removed:"`를 추가하세요.

그다음 `console.log()` 아래에 `console.log()` 변수를 가진 또 다른 `trimmedEnd`를 추가하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

문자열 `"After using the trimEnd() method, trailing spaces removed:"`를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[6], "After using the trimEnd() method, trailing spaces removed:");
```

`trimmedEnd` 변수를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[7], trimmedEnd);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
--fcc-editable-region--

--fcc-editable-region--
```
