---
id: 68f3eeaa65614107ecd5f1f2
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

첫 번째 `console.log()` 아래에 `console.log()`을 하나 더 추가해서 `userInput` 변수를 콘솔에 출력하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`userInput` 변수를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[1], userInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
--fcc-editable-region--

--fcc-editable-region--
```
