---
id: 68f40925d810ae14b2d5a169
title: 16단계
challengeType: 1
dashedName: step-16
---

# --description--

지금 `camelCasedVersion` 변수는 `camelC` 값을 가지고 있습니다. 하지만 원하는 결과는 `camelCase`입니다.

나머지 `ase` 글자를 얻으려면, 음수 인덱스 값을 사용해 `slice()` 메서드를 사용할 수 있습니다. 음수 인덱스 값을 사용하면 `slice()` 메서드가 끝에서부터 추출하도록 지시하는 것임을 기억하세요.

다음은 참고용 예시입니다:

```js
const str = "abcd";
str.slice(-2); // cd
```

`camelCasedVersion` 할당을 업데이트하여 `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()`을 `slice()`에 `lowercaseWord` 메서드를 사용한 결과와도 연결하세요. `-3` 메서드에 숫자 `slice()`을 전달하세요.

이 마지막 변경으로 워크숍이 완료됩니다!

# --hints--

`lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()`과 `lowercaseWord.slice(-3)`을 연결해야 합니다.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase() + lowercaseWord.slice(-3));
```

`"camelCase"`에 문자열 `camelCasedVersion`를 할당하면 안 됩니다. 올바른 방법을 위해 지침을 다시 읽어보세요.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelCase(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

# --solutions--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";

const camelCasedVersion =
  lowercaseWord.slice(0, 5) +
  lowercaseWord[5].toUpperCase() +
  lowercaseWord.slice(-3);

console.log("Camel cased version:");
console.log(camelCasedVersion);
```
