---
id: 68f405c8c9df44132c8fe1bb
title: 15단계
challengeType: 1
dashedName: step-15
---

# --description--

앞서 수업에서 기억했듯이, 카멜 케이스 명명 관례는 첫 단어는 소문자로 시작하고, 그 다음 단어들은 첫 글자가 대문자이고 뒤에 소문자가 옵니다.

다음은 몇 가지 예시입니다:

```md
freeCodeCamp
camelCase
learningIsFun
```

`lowercaseWord` 변수의 두 번째 단어는 `"case"`입니다. 그 단어에서 `c`에 접근하려면 `lowercaseWord[5]`을 사용할 수 있습니다.

`camelCasedVersion` 할당을 업데이트하여 `+` 연산자를 사용해 `lowercaseWord.slice(0, 5)`과 `lowercaseWord[5]`에 대해 문자열을 대문자로 변환하는 올바른 메서드를 사용한 결과를 연결하세요.

이제 콘솔에서 `camelC`이 보여야 합니다.

# --hints--

`lowercaseWord.slice(0, 5)`과 `lowercaseWord[5].toUpperCase()`을 연결해야 합니다.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase());
```

문자열 리터럴 `"camelC"`을 `camelCasedVersion`에 할당하면 안 됩니다. 올바른 방법을 위해 지침을 다시 읽어보세요.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelC(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5)
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
