---
id: 6565a536ba1f9f25bd30e88b
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Projekt implementieren Sie den Luhn-Algorithmus. Dieser Algorithmus ist eine Formel zur Validierung verschiedener Identifikationsnummern.

Beginnen Sie damit, eine Funktion namens `main` zu deklarieren, die als Einstiegspunkt des Programms dient. Verwenden Sie das Schlüsselwort `pass`, um einen Fehler zu vermeiden.

# --hints--

Sie sollten `def main():` in Ihrem Code haben.

```js
({
    test: () => {
        assert(runPython(`
        import inspect
        inspect.isfunction(main)
        `))
    }
})
```

Sie sollten `pass` in Ihrem Code haben.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const pass = __helpers.python.getDef("\n"+transformedCode, "main");
        assert.include(pass?.function_body, "pass");
    }
})
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
