---
id: 65687db6a9fcf3c548a6c1b7
title: Schritt 28
challengeType: 20
dashedName: step-28
---

# --description--

Ein Teil des Algorithmus besteht darin, jede zweite Ziffer von rechts zu verdoppeln. Wenn das Ergebnis der Verdopplung größer oder gleich `10` ist, addieren Sie die beiden Ziffern zusammen. Zum Beispiel, wenn die Ziffer `6` ist, verdoppeln Sie sie, um `12` zu erhalten. Addieren Sie `1` und `2` zusammen, um `3` zu erhalten. Sie können dies tun, indem Sie die _Ganzzahldivision_ verwenden, um die erste Ziffer zu erhalten, und den Modulus-Operator (`%`), um die zweite Ziffer zu erhalten:

```python
my_number = 12
first_digit = my_number // 10
second_digit = my_number % 10
```

Die Ganzzahldivision ergibt den Quotienten der Division, abgerundet auf die nächste ganze Zahl.

Innerhalb der `if`-Anweisung weisen Sie `number` das Ergebnis von `number // 10` (Ganzzahldivision) plus dem Modulus von `number` und `10` zu.

# --hints--

Sie sollten `number = number // 10 + number % 10` innerhalb der `if`-Anweisung haben.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getBlock("\n" + transformedCode, "if number >= 10");
        const { block_body } = verify_card_number;
        const allowedMatches = [
            /number *= *\(?\s*number *\/\/ *10\s*\)? *\+ *\(?\s*number *% *10/,
            /number *= *\(?\s*number *% *10\s*\)? *\+ *\(?\s*number *\/\/ *10/
        ];

        const someMatch = allowedMatches.some(allowedMatch => allowedMatch.test(block_body));
        assert.isTrue(someMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

--fcc-editable-region--
    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            print(number)
            
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
