---
id: 656879a66338b2a461d5d307
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

<!-- TODO: Find better places to split explanation up. -->

Der Luhn-Algorithmus funktioniert wie folgt:

1. Verdoppeln Sie von rechts nach links den Wert jeder zweiten Ziffer; ist das Produkt größer als 9, addieren Sie die Ziffern des Produkts.
2. Addieren Sie alle Ziffern.
3. Wenn die Summe aller Ziffern ein Vielfaches von 10 ist, ist die Nummer gültig; andernfalls ist sie ungültig.

Nehmen Sie als Beispiel eine Kontonummer "7992739871", der eine Prüfziffer hinzugefügt wird, sodass sie die Form 7992739871x hat:

```markdown
Account number      7   9  9  2  7  3  9   8  7  1  x
Double every other  7  18  9  4  7  6  9  16  7  2  x
Sum 2-char digits   7   9  9  4  7  6  9   7  7  2  x
```

Ersetzen Sie die `pass`-Anweisung durch eine Variable `sum_of_odd_digits` mit dem Wert `0`.

# --hints--

Sie sollten `sum_of_odd_digits = 0` innerhalb der `verify_card_number`-Funktion haben.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*=\s*0/);
    }
})
```

Sie sollten `pass` nicht innerhalb der `verify_card_number`-Funktion haben.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.notMatch(function_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    pass

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
