---
id: 656878585631369a6b2d2191
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Das Definieren der Übersetzung übersetzt den String nicht automatisch. Die `translate`-Methode muss auf den zu übersetzenden String mit der Übersetzungstabelle als Argument aufgerufen werden:

```py
my_string = "tamperlot"
translation_table = str.maketrans({'t': 'c', 'l': 'b'})
translated_string = my_string.translate(translation_table)
```

Erstellen Sie eine Variable namens `translated_card_number` und weisen Sie ihr das Ergebnis des Aufrufs der `translate`-Methode auf `card_number` mit `card_translation` als Argument zu.

# --hints--

Sie sollten innerhalb von `translated_card_number` eine `main`-Variable erstellen.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=/);
    }
})
```

Sie sollten `translated_card_number` den Wert `card_number.translate(card_translation)` zuweisen.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=\s*card_number\.translate\(\s*card_translation\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    
--fcc-editable-region--
```
