---
id: 656877c202f44995ed952608
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Ersetzen Sie die `pass`-Anweisung durch eine Variable namens `card_number` mit dem Wert `'4111-1111-4555-1142'`.

# --hints--

Sie sollten `card_number = '4111-1111-4555-1142'` innerhalb der `main`-Funktion haben.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_number\s*=\s*('|")4111-1111-4555-1142\1/);
    }
})
```

Sie sollten keinen `pass` in Ihrem Code haben.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const no_comments = __helpers.python.removeComments(transformedCode);
        assert.notInclude(no_comments, "pass");
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    pass

--fcc-editable-region--
```
