---
id: 65687b48f2201ab32e06c37c
title: Schritt 21
challengeType: 20
dashedName: step-21
---

# --description--

Derzeit wirft Ihr Skript einen `TypeError`, weil Sie versuchen, einen String zu einer Ganzzahl hinzuzufügen. Sie können dies beheben, indem Sie die `digit`-Variable vor dem Hinzufügen zu `sum_of_odd_digits` mit der eingebauten `int`-Funktion in eine Ganzzahl umwandeln:

```python
my_string = '123'
my_int = int(my_string)
```

Wandeln Sie die `digit`-Variable in eine Ganzzahl um, bevor Sie sie zu `sum_of_odd_digits` hinzufügen. Verschieben Sie dann den `print`-Aufruf ans Ende der `verify_card_number`-Funktion, um den Wert von `sum_of_odd_digits` auszugeben.

# --hints--

Sie sollten `sum_of_odd_digits += int(digit)` innerhalb der `for`-Schleife haben.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*\+=\s*int\(\s*digit\s*\)/);
    }
})
```

Sie sollten `print(sum_of_odd_digits)` am Ende der `verify_card_number`-Funktion haben.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /print\(\s*sum_of_odd_digits\s*\)/);
    }
})
```

Sie sollten `print(digit)` nicht mehr haben.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        const no_comments = __helpers.python.removeComments(function_body);

        assert.notMatch(no_comments, /print\(\s*digit\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

--fcc-editable-region--
    for digit in odd_digits:
        print(digit)
        sum_of_odd_digits += digit
        
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
