---
id: 656877f71bba2b97acedf9af
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Python enthält eingebaute Klassen, die uns bei der String-Manipulation helfen können. Eine davon ist die `str`-Klasse. Sie hat eine Methode namens `maketrans`, mit der wir eine Übersetzungstabelle erstellen können. Diese Tabelle kann verwendet werden, um Zeichen in einem String zu ersetzen:

```python
str.maketrans({'t': 'c', 'l': 'b'})
```

Das Obige ersetzt, wenn es auf einen String angewendet wird, alle `t`-Zeichen durch `c` und alle `l`-Zeichen durch `b`.

Erstellen Sie eine Variable namens `card_translation` und weisen Sie ihr eine Übersetzungstabelle zu, die alle `-`- und ` `-Zeichen durch einen leeren String ersetzt.

# --hints--

Sie sollten eine `card_translation`-Variable innerhalb von `main` erstellen.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_translation\s*=/);
    }
})
```

Sie sollten `card_translation` den Wert `str.maketrans({'-': '', ' ': ''})` zuweisen.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        const allowedMatches = [
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|")-\1\s*:\s*('|")\2\s*,\s*('|") \3\s*:\s*('|")\4\s*\}\s*\)/,
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|") \1\s*:\s*('|")\2\s*,\s*('|")-\3\s*:\s*('|")\4\s*\}\s*\)/,
        ];
        const anyMatch = allowedMatches.some((match) => match.test(function_body));
        assert.isTrue(anyMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    
--fcc-editable-region--
```
