---
id: 656879c6f35a85a59c06b3a7
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Sie haben bereits auf Elemente (Zeichen) eines Strings zugegriffen, indem Sie den Indexoperator `[]` verwendet haben. Sie können den Indexoperator auch verwenden, um einen Bereich von Zeichen in einem String mit `string[start:stop:step]` zuzugreifen:

```python
my_string = 'camperbot'
my_string[0:6] == 'camper' # True
my_string[0:6:3] == 'cp' # True
```

Dabei ist `start` der Startindex (einschließlich), `stop` der Endindex (ausschließlich) und `step` die Anzahl der zu überspringenden Zeichen. Wenn nicht angegeben, ist `step` standardmäßig 1.

Erstellen Sie eine Variable namens `card_number_reversed` und weisen Sie ihr die ersten 4 Zeichen von `card_number` zu.

# --hints--

Sie sollten `card_number_reversed = card_number[0:4]` innerhalb der `verify_card_number` Funktion haben.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /card_number_reversed\s*=\s*card_number\s*\[\s*0\s*:\s*4\s*\]/);
    }
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
