---
id: 656b481ec4976439565a78b2
title: Schritt 14
challengeType: 20
dashedName: step-14
---

# --description--

Genauso wie der Schritt optional ist, sind ein Start bei `0` und ein Ende am Ende des Slices optional:

```python
my_string = 'camperbot'
camperbot = my_string[::]
```

Weisen Sie die Umkehrung des gesamten `card_number`-Strings der Variablen `card_number_reversed` zu.

# --hints--

Sie sollten `card_number_reversed = card_number[::-1]` innerhalb der `verify_card_number`-Funktion haben.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /card_number_reversed\s*=\s*card_number\s*\[\s*:\s*:\s*-1\s*\]/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[-1:-5:-1]
    print(card_number_reversed)
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
