---
id: 68e97fe79367ad7b5dd6c9cd
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

In questo passaggio aggiungi un elemento `h1` con il testo `XYZ Bookstore`.

# --hints--

Dovresti avere un elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Il testo del tuo elemento `h1` dovrebbe essere `XYZ Bookstore`. Controlla bene che non ci siano errori di ortografia o maiuscole.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'XYZ Bookstore');
```

Il tuo elemento `h1` dovrebbe essere all'interno dell'elemento `body`.

```js
assert.match(code, /<body>[\s\S]*<h1>[\s\S]*<\/h1>[\s\S]*<\/body>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
