---
id: 68ebdbacdd3fa474132cc975
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

Sotto l'elemento `h1`, aggiungi un elemento `p` con questo testo: `Browse our collection of amazing books!`.

# --hints--

Dovresti avere un elemento `p`.

```js
assert.exists(document.querySelector('p'));
```

Il testo del tuo elemento `p` dovrebbe essere: `Browse our collection of amazing books!`. 

```js
assert.equal(document.querySelector('p')?.innerText.trim(), 'Browse our collection of amazing books!');
```

Il tuo elemento `p` dovrebbe trovarsi sotto il tuo elemento `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
