---
id: 68eca2a795b333ca5fee30a8
title: Passaggio 13
challengeType: 0
dashedName: step-13
---

# --description--

Ora crea una seconda scheda libro. Aggiungi un altro elemento `div` con l'attributo `class` impostato su `card`. Nota come puoi riutilizzare lo stesso nome di classe per più elementi per applicare uno stile coerente.

# --hints--

Dovresti avere due elementi `div` con un attributo `class` che ha valore `card`.

```js
assert.lengthOf(document.querySelectorAll('.card'), 2);
```

Il secondo elemento con classe `card` dovrebbe essere all'interno dell'elemento con classe `card-container`.

```js
assert.lengthOf(document.querySelectorAll('.card-container .card'), 2);
```

I due elementi con classe `card` dovrebbero essere fratelli.

```js
assert.exists(document.querySelector('.card + .card'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
  --fcc-editable-region--
    
  --fcc-editable-region--
  </div>
</body>

</html>
```
