---
id: 68eca11e99e3c5c894ca9d69
title: Passaggio 12
challengeType: 0
dashedName: step-12
---

# --description--

L'elemento `button` viene usato per creare pulsanti cliccabili su una pagina web. I pulsanti sono elementi interattivi su cui gli utenti possono cliccare per eseguire azioni.

Aggiungi un elemento `button` all'interno dell'elemento che ha un `class` di `card`, assegna al pulsante un attributo `class` impostato su `btn` e il testo `Buy Now`.

# --hints--

Dovresti avere un elemento `button` annidato all'interno dell'elemento che ha una classe `card`.

```js
assert.exists(document.querySelector('.card button'));
```

Il tuo elemento `button` dovrebbe avere un attributo `class`.

```js
assert.isTrue(document.querySelector('.card button')?.hasAttribute('class'));
```

Il tuo elemento `button` dovrebbe avere un `class` con valore `btn`.

```js
assert.equal(document.querySelector('.card button')?.className, 'btn');
```

Il testo del tuo elemento `button` dovrebbe essere `Buy Now`.

```js
assert.equal(document.querySelector('.card button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
