---
id: 68eca5c0065c82d256d29ca3
title: Passaggio 17
challengeType: 0
dashedName: step-17
---

# --description--

All'interno della seconda card, aggiungi un elemento `button` con l'attributo `class` impostato su `btn` e il testo `Buy Now`.

Entrambi gli elementi `button` ora condividono lo stesso `class`, il che significa che possono essere stilizzati in modo coerente insieme.

# --hints--

Dovresti avere un elemento `button` annidato all'interno del secondo elemento che ha una classe `card`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('button'));
```

L'elemento `button` della tua seconda card dovrebbe avere un attributo `class`.

```js
const cards = document.querySelectorAll('.card');
assert.isTrue(cards[1]?.querySelector('button')?.hasAttribute('class'));
```

L'elemento `button` della tua seconda card dovrebbe avere un `class` con il valore `btn`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.className, 'btn');
```

Il testo dell'elemento `button` della tua seconda card dovrebbe essere `Buy Now`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
