---
id: 68eca6364c5616d393389a30
title: Passaggio 18
challengeType: 0
dashedName: step-18
---

# --description--

Ricorda, un elemento HTML ha questo aspetto:

```html
<element attribute="value">
    inner text
</element>
```

Sotto l'elemento con la classe `card-container`, aggiungi un nuovo elemento `p` con questo testo:

```md
Review your selections and continue to checkout.
```

Sotto l'elemento `p`, crea un elemento `div` con l'attributo `class` impostato su `btn-container`. Questo contenitore raggrupperà i tuoi elementi di pulsanti di navigazione.

# --hints--

Il nuovo elemento `p` aggiunto dovrebbe trovarsi sotto l'elemento con la classe `card-container`.

```js
assert.exists(document.querySelector('.card-container + p'));
```

Il testo del tuo nuovo elemento `p` dovrebbe essere `Review your selections and continue to checkout.`

```js
const cardContainer = document.querySelector('.card-container');
assert.equal(cardContainer?.nextElementSibling?.innerText.trim(), 'Review your selections and continue to checkout.');
```

Dovresti avere un elemento `div` con un attributo `class` pari a `btn-container`.

```js
assert.exists(document.querySelector('div.btn-container'));
```

Il tuo elemento `div` con un attributo `class` pari a `btn-container` dovrebbe venire dopo il nuovo elemento `p` aggiunto.

```js
assert.exists(document.querySelector('p + .btn-container'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
