---
id: 68ec99e478211dc578699944
title: Passaggio 11
challengeType: 0
dashedName: step-11
---

# --description--

Sotto l'elemento `h2` nel primo elemento con una classe `card`, aggiungi un elemento `p` con il seguente testo:

```md
This is an epic story of Sally and her dog Rex as they navigate through other worlds.
```

# --hints--

Dovresti avere un elemento `p` annidato all'interno dell'elemento con una classe `card`.

```js
assert.exists(document.querySelector('.card p'));
```

Il testo del tuo elemento `p` dovrebbe essere `This is an epic story of Sally and her dog Rex as they navigate through other worlds.`

```js
assert.equal(document.querySelector('.card p')?.innerText.trim(), "This is an epic story of Sally and her dog Rex as they navigate through other worlds.");
```

Il tuo elemento `p` dovrebbe essere sotto il tuo elemento `h2`.

```js
assert.exists(document.querySelector('.card h2 + p'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
