---
id: 68ec98b38d83a3c28dd30efe
title: Passaggio 10
challengeType: 0
dashedName: step-10
---

# --description--

All'interno del primo elemento con classe `card`, aggiungi un elemento `h2` con il testo `Sally's SciFi Adventure`.

# --hints--

Dovresti avere un elemento `h2` annidato all'interno dell'elemento con classe `card`.

```js
assert.exists(document.querySelector('.card h2'));
```

Il testo del tuo elemento `h2` dovrebbe essere `Sally's SciFi Adventure`. 

```js
assert.equal(document.querySelector('.card h2')?.innerText.trim(), "Sally's SciFi Adventure");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
