---
id: 68ec9332a9b5b2b32487bd00
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

L'attributo `id` aggiunge un identificatore univoco a un elemento HTML. Ogni `id` dovrebbe essere unico all'interno di una pagina e dovrebbe essere usato solo una volta.

I valori `id` non possono contenere spazi e dovrebbero contenere solo lettere, cifre, underscore e trattini.

Ecco un esempio:

```html
<p id="para">example paragraph</p>
```

Aggiungi un attributo `id` al tuo elemento con classe `card` e imposta il suo valore su `sally-adventure-book`.

# --hints--

Il tuo elemento con classe `card` dovrebbe avere un attributo `id`.

```js
assert.isTrue(document.querySelector('.card')?.hasAttribute('id'));
```

Il tuo elemento con classe `card` dovrebbe avere un `id` con valore `sally-adventure-book`.

```js
assert.equal(document.querySelector('.card')?.id, 'sally-adventure-book');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
  --fcc-editable-region--
    <div class="card">
  --fcc-editable-region--
      
    </div>
  </div>
</body>

</html>
```
