---
id: 68eca6e5a6759ed4ea0034dc
title: Passaggio 19
challengeType: 0
dashedName: step-19
---

# --description--

All'interno dell'elemento con una classe `btn-container`, aggiungi due elementi `button`:

Primo pulsante:

- Id: `view-cart-btn`
- Classe: `btn`
- Testo: `View Cart`

Secondo pulsante:

- Id: `checkout-btn`
- Classe: `btn`
- Testo: `Checkout`


Congratulazioni! Hai costruito con successo la struttura di una pagina di libreria usando div, classi e id per organizzare il tuo contenuto.

# --hints--

Dovresti avere due elementi `button` annidati dentro l'elemento `.btn-container`.

```js
assert.lengthOf(document.querySelectorAll('.btn-container button'), 2);
```

Il tuo primo pulsante dentro l'elemento con la classe `btn-container` dovrebbe avere un attributo `id` con il valore `view-cart-btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.id, 'view-cart-btn');
```

Il tuo primo pulsante dentro l'elemento con la classe `btn-container` dovrebbe avere un attributo `class` con il valore `btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.className, 'btn');
```

Il testo del tuo primo pulsante dovrebbe essere `View Cart`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.innerText.trim(), 'View Cart');
```

Il tuo secondo pulsante dentro l'elemento con la classe `btn-container` dovrebbe avere un attributo `id` con il valore `checkout-btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.id, 'checkout-btn');
```

Il tuo secondo pulsante dentro l'elemento con `class="btn-container"` dovrebbe avere un attributo `class` con il valore `btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.className, 'btn');
```

Il testo del tuo secondo pulsante dovrebbe essere `Checkout`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.innerText.trim(), 'Checkout');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
  
  <p>Review your selections and continue to checkout.</p>
  
  <div class="btn-container">
  --fcc-editable-region--
    
  --fcc-editable-region--
  </div>
</body>

</html>
```

# --solutions--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>

  <div class="card-container">
     <div class="card" id="sally-adventure-book">
     <h2>Sally's SciFi Adventure</h2>
     <p>
          This is an epic story of Sally and her dog Rex as they navigate
          through other worlds.
     </p>
     <button class="btn">Buy Now</button>
     </div>

     <div class="card" id="dave-cooking-book">
     <h2>Dave's Cooking Adventure</h2>
     <p>
          This is the story of Dave as he learns to cook everything from
          pancakes to pasta, one recipe at a time.
     </p>
     <button class="btn">Buy Now</button>
     </div>
  </div>

  <p>Review your selections and continue to checkout.</p>

  <div class="btn-container">
     <button id="view-cart-btn" class="btn">View Cart</button>
     <button id="checkout-btn" class="btn">Checkout</button>
  </div>
</body>

</html>
```
