---
id: 68ec6cd5b7e8f5a8f7319e32
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

L'elemento `div` viene usato come contenitore per raggruppare altri elementi HTML. Userai principalmente l'elemento `div` quando vuoi raggruppare elementi HTML che condivideranno un insieme di definizioni di stile CSS.

Sotto l'elemento `p`, aggiungi un elemento `div`. Questo `div` sarà un contenitore per le tue schede libro.

**Nota**: questo workshop non applica CSS. Le classi e gli elementi raggruppati sono utili per definire lo stile CSS, ma in questo workshop sono usati solo per strutturare e raggruppare il contenuto. Imparerai come funziona la definizione dello stile in un modulo successivo.

# --hints--

Dovresti avere un elemento `div`.

```js
assert.exists(document.querySelector('div'));
```

Il tuo elemento `div` dovrebbe essere sotto il tuo elemento `p`.

```js
assert.exists(document.querySelector('p + div'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
