---
id: 68ec6d9a315221aa31e54816
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

L'attributo `class` serve per identificare uno o più elementi da stilizzare. A differenza dell'attributo `id`, i nomi delle classi non devono essere unici: più elementi possono condividere la stessa classe.

Ecco un esempio:

```html
<p class="example">example paragraph</p>
```

Aggiungi un attributo `class` al tuo elemento `div` e imposta il suo valore su `card-container`.

# --hints--

Il tuo elemento `div` dovrebbe avere un attributo `class`.

```js
assert.isTrue(document.querySelector('div')?.hasAttribute('class'));
```

Il tuo elemento `div` dovrebbe avere un attributo `class` con il valore `card-container`.

```js
assert.equal(document.querySelector('div')?.className, 'card-container');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  <div>
--fcc-editable-region--
    
  </div>
</body>

</html>
```
