---
id: 68eca5412b5dd9d06b3d6404
title: Passaggio 16
challengeType: 0
dashedName: step-16
---

# --description--

Sotto l'elemento `h2` nella seconda scheda, aggiungi un elemento `p` con questo testo:

```md
This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.
```

# --hints--

Dovresti avere un elemento `p` annidato all'interno del secondo elemento che ha una classe `card`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('p'));
```

Il testo dell'elemento `p` della tua seconda scheda dovrebbe essere `This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.`

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('p')?.innerText.trim(), "This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.");
```

L'elemento `p` della tua seconda scheda dovrebbe trovarsi sotto l'elemento `h2`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2 + p'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
