---
id: 69134d14150d51f8066a3246
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Ora, migliora la struttura del tuo documento HTML per assicurarti che la pagina sia codificata correttamente.

All'interno dell'elemento `head`, aggiungi l'elemento `<meta charset="UTF-8">`.

Infine, aggiungi un elemento `body` sotto la sezione `head`. Qui andrà tutto il contenuto visibile della pagina.

# --hints--

Il tuo codice dovrebbe includere un elemento `<meta charset="UTF-8">`.

```js
assert.exists(document.querySelector('head > meta[charset="utf-8"]'));
```

Il tuo codice dovrebbe includere i tag di apertura e chiusura `<body>`.

```js
assert.match(code,/<body>[\s\S]*<\/body>/i);
```

Il tuo elemento `body` dovrebbe essere sotto l'elemento `head`.

```js
assert.match(code, /<\/head>[\s\S]*<body>[\s\S]*<\/body>[\s\S]*<\/html>/i);
```

# --seed--
## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
  <title>XYZ Bookstore Page</title>
</head>

--fcc-editable-region--
</html>
```
