---
id: 69134b6cb43aa5f254950a10
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo workshop costruirai una pagina per una libreria creando schede di libri che mostrano informazioni su diversi libri. Metterai in pratica l'organizzazione del contenuto usando elementi `div`, classi e ID.

Inizia la tua pagina della libreria creando il codice boilerplate HTML.

Aggiungi la dichiarazione `<!DOCTYPE html>` e gli elementi `html` e `head`.

Aggiungi un attributo `lang` all'elemento `html` e impostalo su `"en"`.

# --hints--

Il tuo codice dovrebbe includere una dichiarazione `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Il tuo codice dovrebbe includere sia i tag di apertura che di chiusura `<html>`.

```js
assert.match(code,/<html\b[^>]*>[\s\S]*<\/html>/i);
```

Il tuo elemento `html` dovrebbe avere un attributo `lang` impostato su `"en"`.

```js
assert.match(code, /<html[^>]*\blang\s*=\s*["']en["'][^>]*>/i);
```

Il tuo codice dovrebbe includere sia i tag di apertura che di chiusura `<head>`.

```js
assert.match(code,/<head>[\s\S]*<\/head>/i);
```

Il tuo elemento `head` dovrebbe essere annidato dentro l'elemento `html`.

```js
assert.match(code, /<\/html>\s*$/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
