---
id: 691f85085787a8c4ca589fbb
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Aggiungi l'elemento `title` all'interno dell'elemento `head`.

Imposta il titolo della pagina su `XYZ Bookstore Page`.

# --hints--

Il tuo codice dovrebbe includere sia i tag di apertura che di chiusura `<title>`.

```js
assert.match(code,/<title>[\s\S]*<\/title>/i);
```

Il tuo titolo dovrebbe essere `XYZ Bookstore Page`.

```js
assert.equal(document.querySelector('title')?.innerText.trim(),'XYZ Bookstore Page');
```

Il tuo elemento `title` dovrebbe essere annidato all'interno dell'elemento `head`.

```js
assert.match(code, /<head>[\s\S]*<title>[\s\S]*<\/title>[\s\S]*<\/head>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
--fcc-editable-region-- 
</head>

</html>
```
