---
id: 65f569725359e10d345bc52a
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

Costruirai un programma che calcola e disegna la traiettoria di un proiettile dato l’angolo, la velocità e l’altezza del lancio.

Inizia importando `math`, lo userai molto in questo progetto perché ha metodi utili come `math.radians`, `math.cos`, `math.sin` e altri.

Crea anche queste variabili per avere il valore dell’accelerazione di gravità e alcuni simboli speciali che ti serviranno più avanti (usa copia e incolla per questi).

```py
GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
```

# --hints--

Usa `import math` per importare il modulo `math`.

```js
({test: () => assert(runPython(`_Node(_code).has_import('import math')`))})
```

Dovresti avere una variabile `GRAVITATIONAL_ACCELERATION` con valore `9.81`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("GRAVITATIONAL_ACCELERATION"), "missing variable"
assert _Node(_code).find_variable("GRAVITATIONAL_ACCELERATION").is_equivalent('GRAVITATIONAL_ACCELERATION = 9.81'), "variable has wrong value"
`)})
```

Dovresti avere una variabile `PROJECTILE` con valore `"∙"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("PROJECTILE"), "missing variable"
assert _Node(_code).find_variable("PROJECTILE").is_equivalent('PROJECTILE = "∙"'), "variable has wrong value"
`)})
```

Dovresti avere una variabile `x_axis_tick` con valore `"T"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("x_axis_tick"), "missing variable"
assert _Node(_code).find_variable("x_axis_tick").is_equivalent('x_axis_tick = "T"'), "variable has wrong value"
`)})
```

Dovresti avere una variabile `y_axis_tick` con valore `"⊣"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("y_axis_tick"), "missing variable"
assert _Node(_code).find_variable("y_axis_tick").is_equivalent('y_axis_tick = "⊣"'), "variable has wrong value"
`)})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
