---
id: 65fd4bd84561a14d3e720061
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

La variabile di classe `__slots__` ha un utilizzo speciale nelle classi Python. Dichiarare `__slots__` e assegnargli una sequenza di stringhe limita la creazione degli attributi a quelli inclusi in quella sequenza. Inoltre, impedisce la creazione dell'attributo speciale `__dict__` e permette un accesso più efficiente agli attributi.

Dovresti usare la variabile `__slots__` all'interno della classe per definire quali attributi ha la classe: assegna a `__slots__` una tupla contenente 3 stringhe, ognuna uguale a uno dei nomi degli attributi definiti in `__init__`.

# --hints--

Dovresti assegnare a `__slots__` una tupla che contiene `'__height'`, `'__speed'` e `'__angle'`.

```js
({test: () => assert(runPython(`
from itertools import permutations
slots = _Node(_code).find_class("Projectile").find_body().find_variable("__slots__")
perms = permutations(("__height", "__speed", "__angle"))
any(slots.is_equivalent(f'__slots__ = {perm}') for perm in perms)
`))})

```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--
class Projectile:

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
--fcc-editable-region--
```
