---
id: 65fd4b3bedee044b4b957d33
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Crea una classe `Projectile` con un metodo `__init__` che inizializza la classe usando tre argomenti: la velocità iniziale, l'altezza iniziale e l'angolo iniziale del lancio del proiettile, in questo ordine.

All'interno del metodo `__init__`, assegna questi argomenti a tre attributi privati: `__speed`, `__height` e `__angle`. L'angolo fornito sarà in gradi; tuttavia, deve essere archiviato internamente in radianti. Per farlo, usa la funzione `math.radians()` per convertire l'angolo da gradi a radianti quando lo assegni a `__angle`.

L'uso di due underscore prima del nome di un attributo attiva il name mangling in Python. Questo significa che gli attributi non sono direttamente consultabili dall'esterno della classe usando i loro nomi dati, e devono essere consultati con i nomi modificati come `ball._Projectile__height` se necessario esternamente, indicando che sono destinati a un uso interno solo.

# --hints--

Dovresti usare `class Projectile:` per creare la classe.

```js
({test: () => assert(runPython(`_Node(_code).has_class("Projectile")`))})

```

Dovresti avere un metodo `__init__` dentro la classe `Projectile`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_body().has_function("__init__")`))})
```

Una nuova istanza della classe `Projectile` dovrebbe avere i valori dei tre attributi correttamente inizializzati.

```js
({test: () => assert(runPython(`
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
p = Projectile(20, 21, 22)
p._Projectile__height == 21 and p._Projectile__speed == 20 and p._Projectile__angle == math.radians(22)
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--

--fcc-editable-region--
```
