---
id: 660400f02cf3a76795b82d4e
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

La prima cosa da impostare è un metodo che calcoli lo spostamento del proiettile, cioè lo spazio orizzontale percorso dal lancio fino a quando il proiettile tocca terra.

Crea un metodo `__calculate_displacement`, che abbia come parametro solo `self`, e restituisca lo spostamento del proiettile.

Usa la seguente formula per calcolare lo spostamento del proiettile:
\\[ d = \frac{v \cdot \cos\theta \cdot \left(v\cdot\sin\theta + \sqrt{v^2 \cdot \sin^2\theta + 2 \cdot g \cdot h}\right)}{g} \\]

Dove $d$ è lo spostamento, $v$ è la velocità iniziale, $\theta$ è l’angolo e $h$ è l’altezza iniziale del proiettile.
Per $g$ puoi usare la variabile `GRAVITATIONAL_ACCELERATION`.

Dovresti usare i metodi `math.cos()` e `math.sin()` per le funzioni trigonometriche e `math.sqrt()` per calcolare la radice quadrata. Inoltre devi sapere che $x^y$ si scrive `x ** y` in python. Inoltre $\sin^2\theta$ significa che il valore risultante dal seno viene poi elevato al quadrato.

Ricorda che con il name mangling devi chiamare il metodo come `_Projectile__calculate_displacement` se vuoi testarlo o usarlo dall’esterno della classe:

```py
ball = Projectile(10, 3, 45)
displacement_of_ball = ball._Projectile__calculate_displacement() # 12.6173996009878
```

# --hints--

Dovresti dichiarare un metodo chiamato `__calculate_displacement` con `def __calculate_displacement(self):`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").has_function("__calculate_displacement")`))})
```

Il metodo `__calculate_displacement` dovrebbe avere come argomento solo `self`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").has_args('self')`))})
```

Il metodo `__calculate_displacement` dovrebbe restituire il valore corretto.

```js
({test: () => assert(runPython(`
p = Projectile(20, 21, 22)
disp = p._Projectile__calculate_displacement()
round(disp, 2) == 55.06 and round(disp, 2) != disp
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
--fcc-editable-region--
    
--fcc-editable-region--
```
