---
id: 660fcf3e1b9bb056b2edb567
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

A questo punto sei pronto per creare la rappresentazione come stringa.
Inizia creando un'istanza della classe `Projectile`. Sotto la definizione della classe, crea una variabile `ball` e assegnale una chiamata a `Projectile` usando `10, 3, 45` come argomenti.

Definisci un metodo `__str__` all'interno della classe, per dare alla tua istanza una corretta rappresentazione come stringa, e fai in modo che restituisca una stringa con il seguente formato:

```py

Projectile details:
speed: 10 m/s
height: 3 m
angle: 45°
displacement: 12.6 m

```

Dovrebbe iniziare e finire con un carattere di nuova riga, l'angolo deve essere scritto come un numero intero in gradi e lo spostamento deve essere stampato con una cifra decimale.

Ti sarà utile `math.degrees` per convertire l'angolo da radianti a gradi.

Quando sei pronto puoi `print(ball)` per testare la tua funzione.

# --hints--

Dovresti avere una variabile `ball`.

```js
({test: () => assert(runPython(`_Node(_code).has_variable('ball')`))})
```

La variabile `ball` dovrebbe avere un valore `Projectile(10, 3, 45)`.

```js
({test: () => assert(runPython(`_Node(_code).find_variable('ball').is_equivalent('ball = Projectile(10, 3, 45)')`))})
```

La rappresentazione come stringa per `Projectile(45, 45, 45)` dovrebbe essere corretta.

```js
({test: () => assert(runPython(
`
ball = Projectile(45, 45, 45)
str(ball) == """
Projectile details:
speed: 45 m/s
height: 45 m
angle: 45°
displacement: 244.4 m
"""`
))})
```

La rappresentazione come stringa dovrebbe essere corretta anche per altre istanze.

```js
({test: () => assert(runPython(`p = Projectile(10, 10, 10)
str(
    p
) == """
Projectile details:
speed: 10 m/s
height: 10 m
angle: 10°
displacement: 15.9 m
"""`))})
```

Dovresti avere una chiamata `print(ball)`.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls('print')
    assert any(c.is_equivalent('print(ball)') for c in calls)
    `)
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)


    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
--fcc-editable-region--

--fcc-editable-region--
        
```
