---
id: 6616ec0bf9ee8548ce6b0f08
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Ora è il momento di lavorare sul calcolo delle coordinate della traiettoria! Crea un metodo della classe `Projectile` chiamato `__calculate_y_coordinate`, che dovrebbe avere, oltre a `self`, un parametro `x`.

\\[ y = y_0 + x \tan\theta - \frac{g x^2}{2 v_0^2 \cos^2\theta} \\]

La formula sopra serve per calcolare la posizione verticale $y$ per una qualsiasi posizione orizzontale $x$, avendo l’angolo iniziale $\theta$, la velocità $v_0$ e l’altezza $y_0$.

Dovrai usare `math.tan()` e `math.cos()` e ricordare che `x ** y` è il modo per scrivere $x^y$, e che il valore di $g$ è nella variabile `GRAVITATIONAL_ACCELERATION`.

Implementa il metodo in modo che restituisca la coordinata $y$.

Se vuoi testare il metodo, dall’esterno della classe puoi chiamare `ball._Projectile__calculate_y_coordinate()` con un numero come argomento dopo che `ball` è stato dichiarato.

# --hints--

La tua classe `Projectile` dovrebbe avere un metodo chiamato `__calculate_y_coordinate`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').has_function('__calculate_y_coordinate')`))}})
```

Il metodo `__calculate_y_coordinate` dovrebbe avere `self, x` come argomenti.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').find_function('__calculate_y_coordinate').has_args('self, x')`))}})
```

Il metodo dovrebbe restituire il valore corretto. Per `Projectile(12, 12, 12)`, date le 4 coordinate $x$ di 0, 1, 2 e 3, la coordinata y dovrebbe essere approssimativamente 12, 12.18, 12.28, 12.32.

```js
({test: () => {
    assert(runPython(`
    a = Projectile(12, 12, 12)
    all([round(a._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(0, 12.0), (1, 12.18), (2, 12.28), (3, 12.32)]])
    `))
}})
```

Il metodo dovrebbe restituire il valore corretto. Per `Projectile(45, 12, 22)`, date le coordinate $x$ di 6, 9, 10 e 12 la coordinata y dovrebbe essere approssimativamente 14.32, 15.41, 15.76, 16.44.

```js
({test: () => {
    assert(runPython(`
    b = Projectile(45, 12, 22)
    all([round(b._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(6, 14.32), (9, 15.41), (10, 15.76), (12, 16.44)]])
    `))
}})
```

Il metodo `__calculate_y_coordinate` non dovrebbe arrotondare il risultato.

```js
({test: () => assert.isFalse(runPython(`Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5) == round(Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5), 2)`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    --fcc-editable-region--
    
    --fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
   
```
