---
id: 662fa453ec8033322b3ab2e3
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

È il momento di parlare di incapsulamento e getter. Hai scritto i tre attributi di istanza come privati usando un doppio underscore iniziale. Nota che questi attributi sono chiamati privati per convenzione: anche se possono ancora essere consultati dall'esterno, si concorda di non farlo.

I getter sono ciò che si può usare per ottenere i valori dall'esterno. Per definire un getter, definisci un metodo che restituisce il valore dell'attributo desiderato e gli assegni un decoratore `@property`:

```py
class Nest:
    ...
    @property
    def number_of_eggs(self):
        return self.__number_of_eggs
```

 Il decoratore trasforma il metodo in una proprietà, il che significa che il metodo non viene chiamato come un metodo normale, ma viene usato come un attributo:

```py
n = Nest()
print(n.number_of_eggs)
```

Nell'esempio sopra, l'attributo privato `__number of eggs` viene consultato tramite la proprietà `number_of_eggs` di `n`.

Crea un getter chiamato `speed` e fallo restituire il valore dell'attributo privato `__speed`.

# --hints--

Dovresti definire un metodo `speed`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').has_function('speed')
`))})
```

Il tuo metodo `speed` dovrebbe avere un solo parametro, `self`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_args('self')
`))})
```

Il tuo metodo `speed` dovrebbe avere un decoratore `@property`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_decorators('property')
`))})
```

Il tuo metodo `speed` dovrebbe restituire il valore dell'attributo privato `__speed`.

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.speed == 22
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

--fcc-editable-region--
    
--fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()
   
```
