---
id: 5d5a813321b9e3db6c106a46
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

JavaScript 是一种强大的语言，它允许你构建交互式网站。

*注意*：对于本课程中的所有剩余项目，你需要具备基本的 HTML 和 CSS 知识。 如果你不熟悉 HTML 和 CSS，请完成[响应式网页设计认证](https: www.freecodecamp.org/learn/2022/responsive-web-design/)。

首先，使用 `DOCTYPE`、`html`、`head` 和 `body` 标签创建你的标准 HTML 模版，然后为 `meta` 添加一个 `charset` 属性。 添加 `title` 元素，将它的文本设置为 `RPG - Dragon Repeller`。 添加一个 `link` 标签，链接到 `styles.css` 文件，以添加样式表。

最后，在你的 `body` 中创建一个 `div` 元素，将其 `id` 设置为 `game`。

# --hints--

你的代码应该包含 `DOCTYPE` 引用。

```js
assert(code.match(/<!DOCTYPE/gi));
```

你应在 `DOCTYPE` 后添加一个空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你应将文档类型定义为 `html` 。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

在类型后，你应该用 `>` 来关闭 `DOCTYPE` 声明。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你的 `html` 元素应该有一个开始标签。 不要忘记设置 `lang` 属性。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素应该有一个结束标签。

```js
assert(code.match(/<\/html\s*>/));
```

你的 `DOCTYPE` 声明应位于 HTML 的开头。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你应该有一个 `head` 开始标签。

```js
assert(code.match(/<head\s*>/i));
```

你应该有一个 `head` 结束标签。

```js
assert(code.match(/<\/head\s*>/i));
```

你应该有一个 `body` 开始标签。

```js
assert(code.match(/<body\s*>/i));
```

你应该有一个 `body` 结束标签。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素应该是同级元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

你的代码应该有一个 `meta` 元素。

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

你的 `meta` 元素应该有一个 `charset` 属性，其值为 `UTF-8`。

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

你的代码应该有一个 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

您应该有一个结尾的 `title` 标记。

```js
assert(code.match(/<\/title\s*>/i));
```

你的代码中应该有一个 `link` 元素。

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

你应该创建一个 `div` 元素。

```js
const div = document.querySelector('div');
assert.exists(div);
```

你的 `div` 元素应该有一个值为 `game` 的 `id`。

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

你的 `div` 元素应该在 `body` 元素内。

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
