---
id: 62a240c67f3dbb1a1e6d95ee
title: 步骤-23
challengeType: 0
dashedName: step-23
---

# --description--

为 `#game` 提供 `500px` 的最大宽度和 `400px` 的最大高度。 将 `background-color` 设置为 `#ffffff`，将 `color` 设置为 `#ffffff`。

使用边距使其居中，将顶部边距设置为 `30px`，底部边距设置为 `0px`，左侧和右侧边距设置为 `auto`。

最后，在其四边添加 `10px` 个填充。

# --hints--

你应该有一个 `#game` 选择器。

```js
const game = new __helpers.CSSHelp(document).getStyle('#game');
assert.exists(game);
```

你的 `#game` 选择器应该具有 `500px` 的 `max-width`。

```js
const maxWidth = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('max-width');
assert.equal(maxWidth, '500px');
```

你的 `#game` 选择器应该具有 `400px` 的 `max-height`。

```js
const maxHeight = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('max-height');
assert.equal(maxHeight, '400px');
```

你的 `#game` 选择器应该具有 `#ffffff` 的 `background-color`。

```js
const background = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('background-color');
assert.equal(background, 'rgb(255, 255, 255)');
```

你的 `#game` 选择器应该具有 `#ffffff` 的 `color`。

```js
const color = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('color');
assert.equal(color, 'rgb(255, 255, 255)');
```

你的 `#game` 选择器应将 `margin-top` 设置为 `30px`。

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-top');
assert.equal(margin, '30px');
```

你的 `#game` 选择器应将 `margin-left` 设置为 `auto`。

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-left');
assert.equal(margin, 'auto');
```

你的 `#game` 选择器应将 `margin-right` 设置为 `auto`。

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-right');
assert.equal(margin, 'auto');
```


你的 `#game` 选择器应将 `margin-bottom` 设置为 `0px`。

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-bottom');
assert.equal(margin, '0px');
```

`#game` 选择器应该有 `10px` 的边距。

```js
const padding = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('padding');
assert.equal(padding, '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
