---
id: 62a257659d0d1e2456f24ba2
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

在开始编写项目代码之前，你应该将其移动到其自己的文件中以保持井然有序。

删除你的 `console.log("Hello World");` 行。 然后将现在为空的 `script` 元素的 `src` 属性设置为 `./script.js`。

# --hints--

你的代码中不应该有 `console.log("Hello World");` 行。

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

你的 `script` 元素应将 `src` 属性设置为 `./script.js`。

```js
const script = document.querySelector("script[data-src$='script.js']");
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>
      console.log("Hello World");
    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
