---
id: 62a3a0a3c0a4b32915d26a6e
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

删除 `console.log("Hello World");` ，然后开始写你自己的项目代码。

在上一个项目中，你学习了如何使用变量和 `let` 关键字，如下所示：

```js
let camperbot;
```

你还学会了如何初始化像这样值的变量：

```js
let age = 32;
```

使用 `let` 关键字声明一个名为 `xp` 的变量，并为其分配数字 `0` 的值。

# --hints--

你的代码中不应该有 `console.log("Hello World");` 行。

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

你应该使用 `let` 关键字声明变量。

```js
assert.match(code, /let/);
```

你应该声明一个名为 `xp` 的变量。

```js
assert.match(code, /let\s+xp/);
```

`xp` 变量值应为 `0`。

```js
assert.equal(xp, 0);
```

你应该将 `xp` 变量初始化为 `0`。

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
