---
id: 62a3a7e4f1060e2fc5ffb34b
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

创建另一个名为 `currentWeaponIndex` 的变量并将其设置为 `0`。

# --hints--

你应该用关键字 `let` 来声明一个名字叫做 `currentWeaponIndex` 的变量。

```js
assert.match(code, /let\s+currentWeaponIndex/i);
```

你应使用驼峰命名法来命名你的变量。

```js
assert.match(code, /currentWeaponIndex/);
```

你的 `currentWeaponIndex` 变量值应设为 `0`。

```js
assert.equal(currentWeaponIndex, 0);
```

你应将你的变量的值初始化为 `0`。

```js
assert.match(code, /let\s+currentWeaponIndex\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
