---
id: 62b46e3a8d4be31be5af793d
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

我们遇到了一点小问题。 你正在尝试在页面上查询按钮元素，但 `script` 标签位于 HTML 的 `head` 中。 这意味着你的代码在浏览器读取 HTML 之前运行，并且你的 `document.querySelector()` 将看不到该按钮 - 因为浏览器尚未处理它。

要解决此问题，请将 `script` 元素移出 `head` 元素，并将其放置在 `body` 元素的末尾（就在结束标签`</body>` 之前）。

# --hints--

你的 `script` 元素不应位于你的 `head` 元素中。

```js
const script = document.querySelector('script[data-src$="script.js"]');
assert.notEqual(script.parentElement.tagName, "HEAD");
```

你的 `script` 元素应位于 `body` 元素的末尾。

```js
const script = document.querySelector('script[data-src$="script.js"]');
const lastChild = document.querySelector('body').lastElementChild;
assert.equal(script, lastChild);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>

  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

let button1 = document.querySelector("#button1");
```
