---
id: 6577ea4f23e3ba1217007bc7
title: 第 28 步
challengeType: 0
dashedName: step-28
---

# --description--

最后，你需要为按钮添加一些样式。 首先将 `cursor` 属性设置为 `pointer`。 然后将文本 `color` 设置为 `#0a0a23`，将 `background-color` 设置为 `#feac32`。

然后将 `background-image` 属性设置为 `linear-gradient(#fecc4c, #ffac33)`。 最后，将 `border` 设置为 `3px solid #feac32`。

# --hints--

你应该有一个 `button` 选择器。

```js
const btn = new __helpers.CSSHelp(document).getStyle('button');
assert.exists(btn);
```

你的 `button` 选择器应将 `cursor` 属性设置为 `pointer`。

```js
const cursor = new __helpers.CSSHelp(document).getStyle('button')?.getPropertyValue('cursor');
assert.equal(cursor, 'pointer');
```

你的 `button` 选择器应将 `color` 属性设置为 `#0a0a23`。

```js
const color = new __helpers.CSSHelp(document).getStyle('button')?.getPropertyValue('color');
assert.equal(color, 'rgb(10, 10, 35)');
```

你的 `button` 选择器应将 `background-color` 属性设置为 `#feac32`。

```js
const backgroundColor = new __helpers.CSSHelp(document).getStyle('button')?.getPropertyValue('background-color');
assert.equal(backgroundColor, 'rgb(254, 172, 50)');
```

你的 `button` 选择器应将 `background-image` 属性设置为 `linear-gradient(#fecc4c, #ffac33)`。

```js
const backgroundImage = new __helpers.CSSHelp(document).getStyle('button')?.getPropertyValue('background-image');
assert.equal(backgroundImage, 'linear-gradient(rgb(254, 204, 76), rgb(255, 172, 51))');
```

你的 `button` 选择器应将 `border` 属性设置为 `3px solid #feac32`。

```js
const border = new __helpers.CSSHelp(document).getStyle('button')?.getPropertyValue('border');
assert.equal(border, '3px solid rgb(254, 172, 50)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
