---
id: 66f26c32ec6f90df01a44f60
title: 연락처 양식 디자인하기
challengeType: 25
dashedName: design-a-contact-form
demoType: onClick
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `div` 클래스를 가진 `form-container`가 있어야 합니다. 이 요소가 양식을 담는 역할을 합니다.

2. `form` 요소 안에 `.form-container` 요소가 있어야 합니다.

3. `form` 요소 안에는 다음 요소들과 입력 필드가 있어야 합니다:

   - 텍스트가 포함된 `h2` 요소.
   - 이름을 위한 텍스트 입력 필드로, `type`가 `text`로 설정되어 있고 `id`, `name`, `required` 속성이 있어야 합니다.
   - 이메일 입력 필드로, `type`가 `email`로 설정되어 있고 `id`, `name`, `required` 속성이 있어야 합니다.
   - 메시지를 위한 텍스트 영역으로, `id`, `name`, 그리고 `required` 속성이 있어야 합니다.
   - 텍스트가 `button`인 `"Submit"` 요소와 `type` 속성이 `submit`로 설정된 요소.

4. 각 `input` 및 `textarea` 요소는 해당 `label` 요소를 가지고 있어야 하며, `for` 속성은 요소의 `id`와 일치해야 합니다.

5. `.form-container` 요소는 배경색이 있어야 하며 `border-radius`, `padding`, `width` 값이 있어야 합니다.

6. `label` 요소들은 여백과 폰트 색상을 가져야 합니다.

7. `input` 및 `textarea` 요소들은 다음 값을 가져야 합니다:
   
   - `width`.
   - `padding`.
   - 아래쪽 여백.

8. `button` 요소는 다음을 가져야 합니다:
   
   - 배경색
   - 기본값과 다른 폰트 크기.

9. 버튼은 배경색이 변경되는 호버 효과가 있어야 합니다.

**참고:** HTML에 스타일시트를 연결하고 CSS를 적용하세요.

# --hints--

`div` 클래스를 가진 `form-container`가 있어야 합니다.

```js
const el = document.querySelector('div.form-container');
assert.isNotNull(el);
```

`form` 요소 안에 `.form-container` 요소가 있어야 합니다.

```js
const el = document.querySelector('.form-container form');
assert.isNotNull(el);
```

`form` 안에는 텍스트가 포함된 `h2` 요소가 있어야 합니다.

```js
const el = document.querySelector('.form-container form h2');
assert.isNotEmpty(el.innerText);
// assert.isAtLeast(document.querySelector('h1').innerText.length, 1);

```

`form` 안에는 타입이 `text`로 설정된 텍스트 입력 필드가 있어야 합니다.

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert.isNotNull(el);
```

`type=text` 입력 필드는 `id`, `name`, `required` 속성을 가져야 합니다.

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert.exists(el);
assert.isTrue(el.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

`form` 안에는 `type`가 `email`로 설정된 이메일 입력 필드가 있어야 합니다.

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert.isNotNull(el);
```

`type=email` 입력 필드는 `id`, `name`, `required` 속성을 가져야 합니다.

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert.isTrue(el?.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

`form` 안에는 `textarea`가 `id`, `name`, `required` 속성을 가져야 합니다.

```js
const el = document.querySelector('.form-container form textarea');
assert.isNotNull(el);
assert.isNotNull(document.querySelector('.form-container form textarea[required]'));
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

`form` 안에는 `button` 요소가 있고 `type` 속성이 `submit`로 설정되어야 합니다.

```js
const el = document.querySelector('.form-container form button');
assert.strictEqual(el?.getAttribute('type'), 'submit');
```

버튼에는 `"Submit"` 텍스트가 있어야 합니다.

```js
const el = document.querySelector('.form-container form button');
assert.strictEqual(document.querySelector('.form-container form button').textContent, 'Submit');
```

`label` 요소에 대해 `#name` 요소가 있어야 하며, `for` 속성은 해당 `id`의 `input` 값과 일치해야 합니다.

```js
const nameInput = document.querySelector('.form-container form input[type="text"]')?.getAttribute('id');

const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

// check if id has a value
assert.isNotEmpty(nameInput);
let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === nameInput) {
        found = true;
        break;
    }
}

assert.isTrue(found);
```

`label` 요소에 대해 `#email` 요소가 있어야 하며, `for` 속성은 해당 `id`의 `input` 값과 일치해야 합니다.

```js
const emailInput = document.querySelector('.form-container form input[type="email"]')?.getAttribute('id');
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);
// check if id has a value
assert.isNotEmpty(emailInput);
let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === emailInput) {
        found = true;
        break;
    }
}

assert.isTrue(found);
```

`label` 요소에 대해 `#message` 요소가 있어야 하며, `for` 속성은 `id`의 `textarea` 값과 일치해야 합니다.

```js
const messageInput = document.querySelector('.form-container form textarea')?.getAttribute('id');
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

// check if id has a value

assert.isNotEmpty(messageInput);

assert.exists(document.querySelector(`label[for="${messageInput}"]`));
```

레이블은 비어 있으면 안 됩니다.

```js
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

for (let label of labels) {
    const labelText = label.textContent.trim();
    assert.isNotEmpty(labelText);
}

```

`.form-container` 요소는 배경색이 있어야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.backgroundColor);
```

`.form-container` 요소는 `border-radius`를 가져야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.borderRadius);
```

`.form-container` 요소는 패딩이 있어야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.padding);
```

`.form-container` 요소는 너비가 있어야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.width);
```

레이블은 여백과 폰트 색상을 가져야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('label')?.margin);
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('label')?.color);
```

`input` 요소들은 너비, 패딩, 아래쪽 여백을 가져야 합니다.

```js
// Helper function to extract properties from multiple selectors
function getStyleProperties(selectors, properties) {
  const result = {};
  properties.forEach((property) => {
    result[property] = selectors.map((selector) => selector?.[property]).find(Boolean);
  });
  return result;
}

// Initializing selectors
const selectors = [
  new __helpers.CSSHelp(document).getStyle('input'),
  new __helpers.CSSHelp(document).getStyle('input, textarea'),
  new __helpers.CSSHelp(document).getStyle('textarea, input')
];

// Extract width, padding, and margin properties
const { width } = getStyleProperties(selectors, ['width']);
assert.isNotEmpty(width);

const {
  padding,
  paddingTop,
  paddingBottom,
  paddingLeft,
  paddingRight
} = getStyleProperties(selectors, ['padding', 'paddingTop', 'paddingBottom', 'paddingLeft', 'paddingRight']);

// Assert padding properties
assert.isOk(
  padding ||
  (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

const { marginBottom } = getStyleProperties(selectors, ['marginBottom']);
assert.isNotEmpty(marginBottom);

```

`textarea` 요소는 너비, 패딩, 아래쪽 여백을 가져야 합니다.

```js
// Helper function to extract properties from multiple selectors
function getStyleProperties(selectors, properties) {
  const result = {};
  properties.forEach((property) => {
    result[property] = selectors.map((selector) => selector?.[property]).find(Boolean);
  });
  return result;
}

// Initializing selectors
const selectors = [
  new __helpers.CSSHelp(document).getStyle('textarea'),
  new __helpers.CSSHelp(document).getStyle('input, textarea'),
  new __helpers.CSSHelp(document).getStyle('textarea, input')
];

// Extract width, padding, and margin properties
const { width } = getStyleProperties(selectors, ['width']);
assert.isNotEmpty(width);

const {
  padding,
  paddingTop,
  paddingBottom,
  paddingLeft,
  paddingRight
} = getStyleProperties(selectors, ['padding', 'paddingTop', 'paddingBottom', 'paddingLeft', 'paddingRight']);

// Assert padding properties
assert.isOk(
  padding ||
  (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

const { marginBottom } = getStyleProperties(selectors, ['marginBottom']);
assert.isNotEmpty(marginBottom);
```

`button` 요소는 배경색이 있어야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button')?.backgroundColor);
```

`button` 요소는 폰트 크기가 있어야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button')?.fontSize);
```

`button` 요소는 배경색이 변경되는 호버 효과가 있어야 합니다.

```js
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
const btn_hover = gs('button:hover')?.backgroundColor;
assert.isNotEmpty(btn_hover);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Contact Form</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Styled Form Project</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="form-container">
        <form>
            <h2>Contact Us</h2>
            <label for="name">Name:</label>
            <input type="text" id="name" name="name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" required>

            <label for="message">Message:</label>
            <textarea id="message" name="message" required></textarea>

            <button type="submit">Submit</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background: linear-gradient(45deg, #83a4d4, #b6fbff);
    font-family: Arial, sans-serif;
}

.form-container {
    background-color: white;
    border-radius: 10px;
    padding: 20px;
    width: 300px;
    text-align: center;
}

form h2 {
    margin-bottom: 20px;
    color: #333;
}

label {
    display: block;
    margin: 10px 0 5px;
    color: #555;
}

input{
    width: 100%;
    padding: 10px;
    margin-bottom: 10px;
    border: 1px solid #ddd;
    border-radius: 5px;
    box-sizing: border-box;
    resize: vertical;
}

textarea {
    width: 100%;
    padding: 10px;
    margin-bottom: 10px;
    border: 1px solid #ddd;
    border-radius: 5px;
    box-sizing: border-box;
    resize: vertical;
}

button {
    background-color: #4CAF50;
    color: white;
    padding: 10px 15px;
    border: none;
    border-radius: 5px;
    cursor: pointer;
    font-size: 1em;
}

button:hover {
    background-color: #45a049;
}
```
