---
id: 6140c7e645d8e905819f1dd4
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Anza na muundo wa msingi wa kawaida. Ongeza tamko lako la `DOCTYPE`, kipengele chako cha `html` chenye lugha imewekwa kuwa Kiingereza, vipengele vyako vya `head` na `body`.

Ongeza kipengele chako cha `meta` kwa `charset` sahihi, kipengele chako cha `title`, na kipengele cha `link` kwa faili ya `./styles.css`.

Weka `title` kuwa `Ferris Wheel`.

# --hints--

Msimbo wako unapaswa kuwa na tamko la `DOCTYPE`.

```js
assert.match(code, /<!DOCTYPE/gi);
```

Unapaswa kuweka nafasi baada ya tamko la `DOCTYPE`.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

Unapaswa kufafanua aina ya hati kuwa `html`.

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

Unapaswa kufunga tamko la `DOCTYPE` kwa lebo ya kufunga `>` baada ya aina.

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

Kipengele chako cha `html` kinapaswa kuwa na lebo ya kufungua yenye sifa ya kipengele `lang` yenye thamani `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Kipengele chako cha `html` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code, /<\/html\s*>/);
```

Tamko lako la `DOCTYPE` linapaswa kuwa mwanzoni mwa HTML yako.

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

Unapaswa kuwa na lebo ya kufungua `head`.

```js
assert.match(code, /<head\s*>/i);
```

Unapaswa kuwa na lebo ya kufunga `head`.

```js
assert.match(code, /<\/head\s*>/i);
```

Unapaswa kuwa na lebo ya kufungua `body`.

```js
assert.match(code, /<body\s*>/i);
```

Unapaswa kuwa na lebo ya kufunga `body`.

```js
assert.match(code, /<\/body\s*>/i);
```

Kipengele chako cha `body` kinapaswa kuwa kipengele cha kwanza baada ya kipengele cha `head`.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

Vipengele vyako vya `head` na `body` vinapaswa kuwa ndugu.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

Kipengele chako cha `head` kinapaswa kuwa kipengele cha kwanza ndani ya kipengele chako cha `html`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

Kipengele chako cha `head` kinapaswa kuwa ndani ya kipengele cha `html`.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Kipengele chako cha `body` kinapaswa kuwa ndani ya kipengele cha `html`.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Msimbo wako unapaswa kuwa na kipengele cha `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Kipengele chako cha `meta` kinapaswa kuwa na sifa ya kipengele `charset` yenye thamani `UTF-8`.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

Kipengele chako cha `meta` kinapaswa kuwa ndani ya kipengele cha `head`.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

Msimbo wako unapaswa kuwa na kipengele cha `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Mradi wako unapaswa kuwa na kichwa cha `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

Kumbuka, herufi kubwa na ndogo pamoja na tahajia ni muhimu kwa kichwa.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

Msimbo wako unapaswa kuwa na kipengele cha `link`.

```js
assert.match(code, /<link/);
```

Unapaswa kuwa na kipengele kimoja cha `link`.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele chako cha `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele `rel` yenye thamani `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele `href` yenye thamani `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
