---
id: 66a8347f97a33e36ffc81b9f
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Para preparar seu código para criar conteúdo, adicione um elemento `body` abaixo do elemento `head`.

# --hints--

Você deve ter uma tag de abertura para `body`.

```js
assert.match(code, /<body>/i);
```

Você deve ter uma tag de fechamento para `body`.

```js
assert.match(code, /<\/body>/i);
```

Você não deve alterar seu elemento `head`. Certifique-se de que não deletou sua tag de fechamento.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

O elemento `body` deve estar depois do elemento `head`.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  
  --fcc-editable-region--
</html>
```
