---
id: 66a83fec026a7a4631e084d2
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Ao trabalhar com elementos `fieldset`, é comum usar uma legenda para descrever o grupo de entradas.Você pode usar o elemento `legend` para isso.

Aqui está um exemplo de como usar um elemento `legend`:

```html
<form action="/example-url">
  <fieldset>
    <legend>Personal Information</legend>
    <!-- inputs go inside here-->
  </fieldset>
</form>
```

Dentro do elemento `fieldset`, adicione um elemento `legend` com o texto `Personal Information`.

# --hints--

Você deve ter uma tag de abertura para `legend`.

```js
assert.match(code, /<legend>/);
```

Você deve ter uma tag de fechamento para `legend`.

```js
assert.match(code, /<\/legend>/);
```

O elemento `legend` deve estar dentro do elemento `fieldset`.

```js
assert.isNotNull(document.querySelector('fieldset legend'));
```

O elemento `legend` deve conter o texto `Personal Information`.

```js
const legend = document.querySelector('fieldset legend');
assert.strictEqual(legend?.innerText.trim(), 'Personal Information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
--fcc-editable-region--
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          
        </fieldset>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>
```
