---
id: 66a84111965a0c46df6bbd0a
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Nas lições anteriores, você aprendeu como associar um elemento `label` com um `input` assim:

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

O atributo `for` no elemento `label` deve corresponder ao atributo `id` no elemento `input`. Isso é conhecido como uma associação explícita.

Abaixo do seu elemento `legend`, adicione um elemento `label` com o texto `Name (required):`. Defina seu atributo `for` com o valor `"full-name"`.

Então, abaixo do seu elemento `label`, adicione um elemento `input` sem atributos. Nos próximos passos, você adicionará os atributos necessários. 

# --hints--

Você deve ter uma tag de abertura para `label`.

```js
assert.match(code, /<label.*>/);
```

Você deve ter uma tag de fechamento para `label`.

```js
assert.match(code, /<\/label>/);
```

Seu elemento `label` deve ser colocado abaixo do seu `legend` dentro do elemento `fieldset`.

```js
assert.isNotNull(document.querySelector('fieldset legend + label'));
```

O elemento `label` deve ter um atributo `for` com o valor `"full-name"`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.getAttribute('for'), 'full-name');
```

Seu elemento `label` deve conter o texto `Name (required):`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.innerText.trim(), 'Name (required):');
```

Você deve ter um elemento `input` abaixo do seu elemento `label`.

```js
assert.isNotNull(document.querySelector('fieldset label + input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
