---
id: 66a83601cd819e37f0dccd14
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Para o texto introdutório, exiba o título principal seguido de uma breve nota sobre como deixar um feedback. 

Dentro do seu elemento `body`, adicione um elemento `header`.

Dentro do elemento `header`, adicione um elemento `h1`.O texto do elemento `h1` deve ser `Hotel Feedback Form`.

Abaixo do elemento `h1`, adicione um elemento `p`.O texto do elemento `p` deve ser `Thank you for staying with us.Please provide feedback on your recent stay.`

# --hints--

Você deve ter uma tag de abertura para `header`.

```js
assert.match(code, /<header>/i);
```

Você deve ter uma tag de fechamento para `header`.

```js
assert.match(code, /<\/header>/i);
```

Você deve acrescentar uma tag de abertura para `h1`.

```js
assert.match(code, /<h1>/i);
```

Você deve acrescentar uma tag de fechamento para `h1`.

```js
assert.match(code, /<\/h1>/i);
```

O elemento `h1` deve ter o texto `Hotel Feedback Form`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.innerText.trim(), 'Hotel Feedback Form');
```

Seu elemento `h1` deve estar dentro do elemento `header`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.parentElement.tagName, 'HEADER');
```

Você deve acrescentar uma tag de abertura para `p`.

```js
assert.match(code, /<p>/i);
```

Você deve acrescentar uma tag de fechamento para `p`.

```js
assert.match(code, /<\/p>/i);
```

O elemento de parágrafo deve ter o texto `Thank you for staying with us.Please provide feedback on your recent stay.`

```js
const pElement = document.querySelector('header p');
assert.strictEqual(pElement?.innerText.trim(), 'Thank you for staying with us. Please provide feedback on your recent stay.');
```

O elemento de parágrafo deve estar dentro do elemento `header`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.parentElement.tagName, 'HEADER');
```

Seu elemento de parágrafo deve vir depois do seu `h1`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.previousElementSibling.tagName, 'H1');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
