---
id: 66a83bdcf425e7446900b7c4
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Nas lições anteriores, você aprendeu como trabalhar com o elemento `form` assim:

```html
<form method="value-goes-here" action="url-goes-here">
  <!-- inputs go inside here -->
</form>
```

O atributo `action` é usado para especificar para onde os dados do formulário devem ser enviados quando o formulário é enviado.

O atributo `method` é usado para especificar o método HTTP a ser usado ao enviar os dados do formulário. Os métodos mais comuns são `GET` e `POST`. 

**NOTA**: Você aprenderá sobre como os métodos HTTP funcionam em módulos posteriores.

Dentro do elemento `main`, adicione um elemento `form` com um atributo `action` definido como `"https://hotel-feedback.freecodecamp.org"` e um atributo `method` definido como `"POST"`.

# --hints--

Você deve ter uma tag de abertura para `form`.

```js
assert.match(code, /<form.*>/i);
```

Você deve ter uma tag de fechamento para `form`.

```js
assert.match(code, /<\/form>/i);
```

O elemento `form` deve ter um atributo `action` definido como `"https://hotel-feedback.freecodecamp.org"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('action'), 'https://hotel-feedback.freecodecamp.org');
```

O elemento `form` deve ter um atributo `method` definido como `"POST"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('method'), 'POST');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
