---
id: 66a83e5e491625454b6f62c3
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Formulários consistem em `inputs` onde os usuários podem inserir seus dados. Você pode agrupar entradas relacionadas usando o elemento `fieldset`. 

Aqui está um exemplo de uso de um elemento `fieldset`:

```html
<form action="/example-url">
  <fieldset>
  <!-- inputs go inside here-->
  </fieldset>
</form>
```

Dentro do elemento `form`, adicione um elemento `fieldset`.

# --hints--

Você deve ter uma tag de abertura para `fieldset`.

```js
assert.match(code, /<fieldset>/i);
```

Você deve ter uma tag de fechamento para `fieldset`.

```js
assert.match(code, /<\/fieldset>/i);
```

O elemento `fieldset` deve estar dentro do elemento `form`.

```js
assert.isNotNull(document.querySelector('form fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">        
--fcc-editable-region--
        
--fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
