---
id: 66a9577022877d72d8f43b4f
title: Passo 21
challengeType: 0
dashedName: step-21
---

# --description--

A próxima seção do formulário perguntará aos usuários por que escolheram aquele hotel específico. Os usuários terão a oportunidade de selecionar várias opções.

Comece adicionando outro elemento `fieldset`.

Dentro do elemento `fieldset`, adicione um elemento `legend` com o texto `Why did you choose to stay at our hotel? (Check all that apply)`.

# --hints--

You should have a `fieldset` element.

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 3);
```

Você deve ter um elemento `legend` dentro do elemento `fieldset`.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(3) legend'), 1);
```

O elemento `legend` deve ter o texto de `Why did you choose to stay at our hotel? (Check all that apply)`. Verifique novamente erros ortográficos e problemas de espaçamento.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) legend')?.innerText, 'Why did you choose to stay at our hotel? (Check all that apply)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        --fcc-editable-region--

        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
