---
id: 66a9419e2d18476c645ce693
title: Passo 14
challengeType: 0
dashedName: step-14
---

# --description--

O `input` de números é usado para criar um campo de entrada numérico. 

Aqui está um exemplo de um campo de entrada numérico:

```html
<input type="number" id="age" name="age" min="18" max="100">
```

Os atributos `min` e `max` são usados para definir os valores mínimos e máximos que podem ser inseridos no campo de entrada.

Abaixo do elemento `label`, adicione uma entrada com o atributo `type` definido como `"number"` e um `id` de `"age"`.

O atributo `name` deve ser definido como `"age"`, o atributo `min` deve ser definido como `"3"` e o atributo `max` deve ser definido como `"100"`.

# --hints--

Você deve ter um `input` numérico abaixo do seu elemento `label`.

```js
assert.isNotNull(document.querySelector('label + input[type="number"]'));
```

O `input` de números deve ter um `id` de `age`.

```js
assert.isNotNull(document.querySelector('label + input#age'));
```

O `input` de números deve ter um atributo `name` definido como `age`.

```js
assert.isNotNull(document.querySelector('label + input[name="age"]'));
```

O `input` de números deve ter um atributo `min` definido como `"3"`.

```js
assert.isNotNull(document.querySelector('label + input[min="3"]'));
```

O `input` de números deve ter um atributo `max` definido como `"100"`.

```js
assert.isNotNull(document.querySelector('label + input[max="100"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
      --fcc-editable-region--
          
      --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
