---
id: 66a9521bc70162712caf118d
title: Passo 16
challengeType: 0
dashedName: step-16
---

# --description--

Se você deseja que os usuários selecionem uma opção de uma lista de opções, você pode usar um conjunto de botões de opção (ou radio buttons, em inglês).

Aqui vemos um exemplo de dois botões de seleção:

```html
<input type="radio" id="yes" name="first-time">
<label for="yes">Yes</label>
<input type="radio" id="no" name="first-time">
<label for="no">No</label>
```

Neste exemplo, os botões de seleção são agrupados usando o mesmo valor do atributo `name`. Isso significa que apenas um botão de opção pode ser selecionado por vez.

Abaixo do seu elemento `legend`, adicione um botão `radio` com o `id` definido como `"yes-option"`, o atributo `name` definido como `"hotel-stay"` e um atributo `value` definido como `"yes"`.

Abaixo do seu botão `radio`, adicione um elemento `label` com o texto `Yes` e um atributo `for` definido como `"yes-option"`.

# --hints--

Você deve ter um botão `radio` abaixo do seu elemento `legend`.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]'));
```

O botão `radio` deve ter um atributo `id` definido como `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.id, 'yes-option');
```

O botão `radio` deve ter um atributo `name` definido como `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.name, 'hotel-stay');
```

Seu botão de opção deve ter um atributo `value` definido como `"yes"`.

```js
assert.strictEqual(
  document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.value,
  'yes'
);
```

Você deve ter um elemento `label` abaixo do seu botão `radio`.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label'));
```

O elemento `label` deve ter o atributo `for` definido como `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.htmlFor, 'yes-option');
```

O elemento `label` deve ter o texto `Yes`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.innerText?.trim(), 'Yes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
