---
id: 66a937e74920ba68ebe5e86d
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

Nas lições anteriores, você aprendeu como trabalhar com os atributos `placeholder` e `required` assim:

```html
<input type="text" placeholder="Ex. John Doe" required>
```

Para o elemento `input` existente, adicione um atributo `placeholder` com o valor de `"Ex. John Doe"`.

Além disso, adicione o atributo `required` ao elemento `input`.

# --hints--

Você não deve remover o elemento `input` existente.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

O `input` deve ter um atributo `placeholder` com o valor de `"Ex. John Doe"`.

```js
assert.strictEqual(document.querySelector('fieldset input')?.getAttribute('placeholder'), 'Ex. John Doe');
```

O elemento `input` deve ter um atributo `required`.

```js
assert.isNotNull(document.querySelector('fieldset input[required]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
