---
id: 66a962954f4e0d76223b37ed
title: Passo 21
challengeType: 0
dashedName: step-21
---

# --description--

Em seguida, adicione outro `input` do tipo checkbox com os atributos `id`, `name` e `value` definidos como `"location"`.

Para o elemento `label`, o texto de `Location` e o atributo `for` devem ser definidos como `"location"`.

Abaixo daquele elemento `label`, adicione outro `input` do tipo checkbox com os atributos `id`, `name` e `value` definidos como `"reputation"`.

Para o elemento `label`, o texto de `Reputation` e o atributo `for` devem ser definidos como `"reputation"`.

# --hints--

Você deve ter um terceiro `input` de caixa de seleção.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)"));
```

Sua terceira caixa de seleção `input` deve ter um atributo `id` definido como `"location"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('id'), 'location');
```

Sua terceira caixa de seleção `input` deve ter um atributo `name` definido como `"location"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('name'), 'location');
```

Sua terceira caixa de seleção `input` deve ter um atributo `value` definido como `"location"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('value'), 'location');
```

Você deve ter um elemento `label` após seu `input` de checkbox com o texto `Location`.

```js 
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[value='location'] + label")?.textContent?.trim(), 'Location');
```

O elemento `label` deve ter o atributo `for` definido como `"location"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="location"] + label')?.getAttribute('for'), 'location');
```

Você deve ter uma quarta caixa de seleção `input` abaixo do seu `label`.

```js 
assert.exists(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']"));
```

Você deve ter uma quarta caixa de seleção `input` com um atributo `id` definido como `"reputation"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('id'), 'reputation');
```

Sua quarta caixa de seleção `input` deve ter um atributo `name` definido como `"reputation"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('name'), 'reputation');
```

Sua quarta caixa de seleção `input` deve ter um atributo `value` definido como `"reputation"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('value'), 'reputation');
```

Você deve ter um elemento `label` com o texto `Reputation` abaixo da caixa de seleção de reputação.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.textContent?.trim(), 'Reputation');
```

O elemento `label` deve ter o atributo `for` definido como `"reputation"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.getAttribute('for'), 'reputation');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="ads" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="recommendation"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
