---
id: 66a9689b1bf24b7750898a1b
title: Passo 23
challengeType: 0
dashedName: step-23
---

# --description--

Para o `input` final e o `label` dentro deste `fieldset`, você adicionará um `input` do tipo checkbox com os atributos `id` e `value` definidos como `"price"` e um atributo `name` definido como `"choice"`.

Então, um elemento `label` com o texto `Price` e o atributo `for` definido como `"price"`.

Agora você pode testar o `form` selecionando as várias caixas de seleção.

# --hints--

Você deve ter um quinto `input` de caixa de seleção.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(3) input[type="checkbox"]:nth-of-type(5)'));
```

Sua quinta caixa de seleção `input` deve ter um atributo `id` definido como `"price"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(5)")?.getAttribute('id'), 'price');
```

Sua quinta caixa de seleção `input` deve ter um atributo `name` definido como `"choice"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(5)")?.getAttribute('name'), 'choice');
```

Sua quinta caixa de seleção `input` deve ter um atributo `value` definido como `"price"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(5)")?.getAttribute('value'), 'price');
```

Você deve ter um elemento `label` após seu `input` de checkbox com o texto `Price`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="price"] + label')?.textContent?.trim(), 'Price');
```

O elemento `label` deve ter o atributo `for` definido como `"price"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="price"] + label')?.getAttribute('for'), 'price');
```   

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          <input type="checkbox" id="location" name="choice" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="choice"
            value="reputation"
          />
          <label for="reputation">Reputation</label>
          
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
