---
id: 66a954b2bcddba72076c1857
title: Passo 20
challengeType: 0
dashedName: step-20
---

# --description--

Abaixo do seu elemento `label`, adicione um botão `radio` com o `id` definido como `"no-option"` e o atributo name definido como `"hotel-stay"`.

Abaixo do seu novo botão `radio`, adicione outro elemento `label` com o atributo `for` definido como `"no-option"`. O texto para o `label` deve ser `No`.

Quando você terminar, poderá testar os botões de opção selecionando uma opção por vez.

# --hints--

Você deve ter um elemento `input` abaixo do seu primeiro elemento `label`.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) label + input'));
```

Seu elemento `input` deve ter um atributo `type` com valor `radio`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.getAttribute('type'), 'radio');
```

Seu elemento `input` deve ter um atributo `id` com o valor `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.id, 'no-option');
```

Seu elemento `input` deve ter um atributo `name` com o valor `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.name, 'hotel-stay');
```

Você deve ter um segundo elemento `label` abaixo do seu segundo elemento `input`.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label'));
```

Seu elemento `label` deve ter um atributo `for` definido como `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.htmlFor, "no-option");
```

Seu elemento `label` deve ter o texto `No`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.textContent?.trim(), 'No');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          --fcc-editable-region--
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
