---
id: 66a93bbe65a26169dbf3bc39
title: Passo 13
challengeType: 0
dashedName: step-13
---

# --description--

O formulário de feedback de hotéis deve também coletar um endereço de e-mail do usuário. 

Comece adicionando um novo elemento `label` com o texto `Email address (required):` ao formulário.O elemento `label` deve ter o atributo `for` definido com o valor `"email"`.

# --hints--

Você deve ter um elemento `label`.

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 2);
```

Seu elemento `label` deve conter o texto `Email address (required):`.

```js
assert.strictEqual(document.querySelector('fieldset label:nth-of-type(2)')?.innerText.trim(), 'Email address (required):');
```

O elemento `label` deve ter o atributo `for` definido com o valor `"email"`.

```js
assert.strictEqual(document.querySelector('fieldset label[for="email"]')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
