---
id: 66a93730719e1f68410cce54
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

O atributo `name` é usado para identificar os dados do formulário depois que ele foi enviado para o servidor.

Aqui está um exemplo de como usar o atributo `name`:

```html
<input type="email" name="email">
```

Adicione um atributo `name` ao elemento `input` com o valor de `"name"`.

# --hints--

Você não deve remover o elemento `input` existente.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Você não deve criar um elemento `input`.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

O elemento `input` deve ter um atributo `name` com o valor de `"name"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('name'), 'name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
