---
id: 67a51d3a8a6fe123b77b0c6e
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

Elementos `input` podem ter um atributo `size`. Este atributo define o número de caracteres que devem estar visíveis enquanto o usuário digita no input. O valor de `size` deve ser um inteiro não negativo maior que zero. Se `size` não for especificado ou for especificado com um valor inválido, a entrada terá a largura padrão definida pelo navegador.

```html
<label for="lastName">Last Name:</label>
<input id="lastName" name="lastName" type="text" size="10" />
```

Dê aos inputs de nome e email um atributo `size` com valor `"20"`.

# --hints--

Você deve dar ao nome `input` um atributo `size` com o valor `"20"`.

```js
const label1 = document.querySelector('label:nth-of-type(1) + input');
assert.strictEqual(label1?.getAttribute('size'), '20');
```

Você deve dar ao `input` de email um atributo `size` com o valor `"20"`.

```js
const label2 = document.querySelector('label:nth-of-type(2) + input');
assert.strictEqual(label2?.getAttribute('size'), '20');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />         
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
