---
id: 657a0ea50da0c8d9d6d7950a
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在這個項目中，你將通過構建一個骰子游戲來學習算法思維。 總共有 6 輪，每輪玩家最多可以投擲骰子 3 次並獲得分數。

HTML 和 CSS 已爲你提供。 請隨意探索它們。 準備就緒後，你將需要設置 HTML 變量。 獲取所有 `.die` 元素並將它們分配給 `listOfAllDice` 變量。 獲取你的分數輸入（`#score-options` 中的 `input` 元素）和分數跨度，並將它們分配給 `scoreInputs` 和 `scoreSpans`。 將 `#current-round` 元素分配給 `roundElement`，並將 `#current-round-rolls` 元素分配給 `rollsElement`，然後對 `#total-score` 和 `#score-history` 元素執行相同操作。 將 `#roll-dice-btn`、`#keep-score-btn`、`#rules-btn` 和 `.rules-container` 分配給具有正確格式名稱的變量。

當用戶點擊 `Show rules` 按鈕時，他們應該能夠在顯示和隱藏遊戲規則之間切換。 創建一個變量 `isModalShowing` 來跟蹤遊戲規則的狀態。

每次用戶擲骰子時，你都需要跟蹤所有骰子的值。 創建一個變量 `diceValuesArr` 來跟蹤這一點。

在整個遊戲過程中，你需要跟蹤當前得分、總得分、擲骰子次數以及玩家所處的回合。 爲此目的，聲明 `rolls`、`score` 和 `round` 變量。

思考一下每個變量的起始值應該是多少。 所有這些變量都應該能夠重新分配。

# --hints--

你應該將 `.die` 元素分配給 `listOfAllDice`。

```js
assert.deepEqual(listOfAllDice, document.querySelectorAll(".die"));
```

你應該將 `#score-options input` 元素分配給 `scoreInputs`。

```js
assert.deepEqual(scoreInputs, document.querySelectorAll("#score-options input"));
```

你應該將 `#score-options span` 元素分配給 `scoreSpans`。

```js
assert.deepEqual(scoreSpans, document.querySelectorAll("#score-options span"));
```

你應該將 `#current-round` 元素分配給 `roundElement`。

```js
assert.deepInclude([document.getElementById("current-round"), document.querySelector("#current-round")], roundElement);
```

你應該將 `#current-round-rolls` 元素分配給 `rollsElement`。

```js
assert.deepInclude([document.getElementById("current-round-rolls"), document.querySelector("#current-round-rolls")], rollsElement);
```

你應該將 `#total-score` 元素分配給 `totalScoreElement`。

```js
assert.deepInclude([document.getElementById("total-score"), document.querySelector("#total-score")], totalScoreElement);
```

你應該將 `#score-history` 元素分配給 `scoreHistory`。

```js
assert.deepInclude([document.getElementById("score-history"), document.querySelector("#score-history")], scoreHistory);
```

你應該將 `#roll-dice-btn` 元素分配給 `rollDiceBtn`。

```js
assert.deepInclude([document.getElementById("roll-dice-btn"), document.querySelector("#roll-dice-btn")], rollDiceBtn);
```

你應該將 `#keep-score-btn` 元素分配給 `keepScoreBtn`。

```js
assert.deepInclude([document.getElementById("keep-score-btn"), document.querySelector("#keep-score-btn")], keepScoreBtn);
```

你應該將 `#rules-btn` 元素分配給 `rulesBtn`。

```js
assert.deepInclude([document.getElementById("rules-btn"), document.querySelector("#rules-btn")], rulesBtn);
```

你應該將 `.rules-container` 元素分配給 `rulesContainer`。

```js
assert.deepEqual(rulesContainer, document.querySelector(".rules-container"));
```

你的 `isModalShowing` 變量應該具有值 `false`。

```js
assert.isFalse(isModalShowing);
```

你的 `diceValuesArr` 應該有一個空數組作爲值。

```js
assert.isArray(diceValuesArr);
assert.isEmpty(diceValuesArr);
```

你的 `rolls` 變量的值應該爲 `0`。

```js
assert.strictEqual(rolls, 0);
```

你的 `score` 變量的值應該爲 `0`。

```js
assert.strictEqual(score, 0);
```

你的 `round` 變量應該具有值 `1`。

```js
assert.strictEqual(round, 1);
```

所有新變量都應該使用 `let` 聲明。

```js
assert.match(code, /let isModalShowing/);
assert.match(code, /let diceValuesArr/);
assert.match(code, /let rolls/);
assert.match(code, /let score/);
assert.match(code, /let round/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Advanced Dice Game</title>
  <link rel="stylesheet" href="styles.css" />
</head>

<body>
  <header>
    <h1>Advanced Dice Game</h1>
    <button class="btn" id="rules-btn" type="button">Show rules</button>
    <div class="rules-container">
      <h2>Rules</h2>
      <ul>
        <li>There are total of six rounds</li>
        <li>You can only roll the dice three times per round</li>
        <li>To start the game, roll the dice</li>
        <li>
          Then, choose from one of the selected scores or roll the dice again
        </li>
        <li>
          If you choose a selected score, then you will move to the next round
        </li>
        <li>
          If you decline to choose a selected score, then you can roll the
          dice again two more times
        </li>
      </ul>
      <h2 class="points">Points</h2>
      <ul>
        <li>Three of a kind: Sum of all five dice</li>
        <li>Four of a kind: Sum of all five dice</li>
        <li>Full house: Three of a kind and a pair - 25 points</li>
        <li>
          Small straight: Four of the dice have consecutive values - 30 points
        </li>
        <li>
          Large straight: All five dice have consecutive values - 40 points
        </li>
      </ul>
    </div>
  </header>

  <main>
    <form id="game">
      <div id="dice">
        <div class="die"></div>
        <div class="die"></div>
        <div class="die"></div>
        <div class="die"></div>
        <div class="die"></div>
      </div>

      <p class="rounds-text">
        <strong>Rolls:</strong> <span id="current-round-rolls">0</span> |
        <strong>Round:</strong> <span id="current-round">1</span>
      </p>

      <div id="score-options">
        <div>
          <input type="radio" name="score-options" id="three-of-a-kind" value="three-of-a-kind" disabled />
          <label for="three-of-a-kind">Three of a kind<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="four-of-a-kind" value="four-of-a-kind" disabled />
          <label for="four-of-a-kind">Four of a kind<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="full-house" value="full-house" disabled />
          <label for="full-house">Full house<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="small-straight" value="small-straight" disabled />
          <label for="small-straight">Small straight<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="large-straight" value="large-straight" disabled />
          <label for="large-straight">Large straight<span></span></label>
        </div>

        <div>
          <input type="radio" name="score-options" id="none" value="none" disabled />
          <label for="none">None of the above<span></span></label>
        </div>
      </div>

      <button class="btn" id="keep-score-btn" type="button">
        Keep the above selected score
      </button>
      <button class="btn" id="roll-dice-btn" type="button">
        Roll the dice
      </button>
    </form>

    <div id="scores">
      <h3>Score history (Total score: <span id="total-score">0</span>)</h3>
      <ol id="score-history"></ol>
    </div>
  </main>
  <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --black: #000;
  --white: #fff;
  --grey: #3b3b4f;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

header {
  color: var(--light-grey);
  text-align: center;
}

h1 {
  font-size: 2.5rem;
  margin: 25px 0;
}

.rules-container {
  display: none;
  background-color: var(--light-grey);
  color: var(--black);
  width: 50%;
  margin: 20px auto;
  height: 300px;
  border-radius: 10px;
  overflow-y: scroll;
}

.rules-container ul {
  list-style-type: none;
}

.points {
  margin-top: 15px;
}

main {
  background-color: var(--light-grey);
  padding: 20px;
  display: grid;
  grid-template-columns: repeat(auto-fit, minmax(500px, 1fr));
  gap: 1rem;
  margin: auto;
  justify-items: center;
  width: 50%;
  border-radius: 10px;
}

#dice {
  display: flex;
  justify-content: space-around;
  margin-bottom: 15px;
}

.die {
  width: 40px;
  height: 40px;
  text-align: center;
  margin-right: 15px;
  border: 4px solid var(--black);
  padding: 10px;
}

.rounds-text {
  text-align: center;
}

input[type="radio"]:disabled + label {
  color: var(--grey);
}

#score-history {
  margin-top: 15px;
  text-align: center;
  list-style-position: inside;
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px 0 10px 0.5rem;
  color: var(--black);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover:enabled {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

@media (max-width: 992px) {
  main {
    width: 100%;
  }
}

@media (max-width: 500px) {
  .btn {
    width: 120px;
  }
}

```

```js
--fcc-editable-region--

--fcc-editable-region--
```
