---
id: 587d7fad367417b2b2512be1
title: Klick-Ereignisse mit JavaScript über die Eigenschaft onclick behandeln
challengeType: 6
forumTopicId: 301503
dashedName: handle-click-events-with-javascript-using-the-onclick-property
---

# --description--

Sie möchten, dass Ihr Code erst ausgeführt wird, wenn Ihre Seite vollständig geladen ist. Dafür können Sie ein JavaScript-Ereignis namens `DOMContentLoaded` an das Dokument anhängen. Hier ist der Code, der das macht:

```js
document.addEventListener('DOMContentLoaded', function () {});
```

Sie können Handler implementieren, die innerhalb der `DOMContentLoaded`-Funktion liegen. Sie können einen `onclick`-Ereignis-Handler implementieren, der ausgelöst wird, wenn der Benutzer auf das `#getMessage`-Element klickt, indem Sie den folgenden Code hinzufügen:

```js
document.getElementById('getMessage').onclick = function () {};
```

# --instructions--

Fügen Sie innerhalb der `DOMContentLoaded`-Funktion einen Klick-Ereignis-Handler für das Element mit der ID `getMessage` hinzu.

# --hints--

Ihr Code sollte die `document.getElementById`-Methode verwenden, um das Element mit der ID `getMessage` auszuwählen.

```js
assert.match(code, /document\s*\.getElementById\(\s*('|")getMessage\1\s*\)/g);
```

Ihr Code sollte einen `onclick`-Ereignis-Handler hinzufügen.

```js
assert.isFunction(document.getElementById('getMessage').onclick);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    // Add your code below this line
    // Add your code above this line
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    // Add your code below this line
    document.getElementById('getMessage').onclick = function () {};
    // Add your code above this line
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
