---
id: 587d7faf367417b2b2512be8
title: "Geolokalisierungsdaten abrufen, um die GPS-Koordinaten eines Benutzers zu finden"
challengeType: 6
forumTopicId: 18188
dashedName: get-geolocation-data-to-find-a-users-gps-coordinates
---

# --description--

Eine weitere coole Sache, die Sie tun können, ist, auf den aktuellen Standort Ihres Benutzers zuzugreifen. Jeder Browser hat einen eingebauten Navigator, der Ihnen diese Informationen geben kann.

Der Navigator ermittelt die aktuelle Länge und Breite des Benutzers.

Sie werden aufgefordert, zu erlauben oder zu blockieren, dass diese Seite Ihren aktuellen Standort kennt. Die Aufgabe kann in beiden Fällen abgeschlossen werden, solange der Code korrekt ist.

Wenn Sie erlauben auswählen, sehen Sie, wie sich der Text auf dem Ausgabegerät in Ihre Breiten- und Längengrade ändert.

Hier ist der Code, der dies macht:

```js
if (navigator.geolocation) {
  navigator.geolocation.getCurrentPosition(function (position) {
    document.getElementById('data').innerHTML =
      'latitude: ' +
      position.coords.latitude +
      '<br>longitude: ' +
      position.coords.longitude;
  });
}
```

Zuerst wird geprüft, ob das `navigator.geolocation`-Objekt existiert. Falls ja, wird die `getCurrentPosition`-Methode dieses Objekts aufgerufen, die eine asynchrone Anfrage für die Position des Benutzers startet. Wenn die Anfrage erfolgreich ist, läuft die Callback-Funktion in der Methode. Diese Funktion greift mit Punktnotation auf die Werte für Breite und Länge des `position`-Objekts zu und aktualisiert das HTML.

# --instructions--

Fügen Sie den Beispielcode innerhalb der `script`-Tags ein, um den aktuellen Standort eines Benutzers zu überprüfen und ihn in das HTML einzufügen.

# --hints--

Ihr Code sollte `navigator.geolocation` verwenden, um auf den aktuellen Standort des Benutzers zuzugreifen.

```js
assert.match(code, /navigator\.geolocation\.getCurrentPosition/g);
```

Ihr Code sollte `position.coords.latitude` verwenden, um den Breitengrad des Benutzers anzuzeigen.

```js
assert.match(code, /position\.coords\.latitude/g);
```

Ihr Code sollte `position.coords.longitude` verwenden, um den Längengrad des Benutzers anzuzeigen.

```js
assert.match(code, /position\.coords\.longitude/g);
```

Sie sollten die Position des Benutzers innerhalb des `div`-Elements mit `id="data"` anzeigen.

```js
assert.match(
  code,
  /document\.getElementById\(\s*('|")data\1\s*\)\.innerHTML/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  // Add your code below this line

  // Add your code above this line
</script>
<h4>You are here:</h4>
<div id="data"></div>
```

# --solutions--

```html
<script>
  // Add your code below this line
  if (navigator.geolocation) {
    navigator.geolocation.getCurrentPosition(function(position) {
      document.getElementById('data').innerHTML = "latitude: " + position.coords.latitude + "<br>longitude: " + position.coords.longitude;
    });
  }
  // Add your code above this line
</script>
<h4>You are here:</h4>
<div id="data">

</div>

</section>
```
