---
id: 587d7fae367417b2b2512be3
title: JSON mit der JavaScript XMLHttpRequest-Methode abrufen
challengeType: 6
forumTopicId: 301502
dashedName: get-json-with-the-javascript-xmlhttprequest-method
---

# --description--

Sie können auch Daten von einer externen Quelle anfordern. Hier kommen APIs ins Spiel.

Denken Sie daran, dass APIs – oder Application Programming Interfaces – Werkzeuge sind, die Computer verwenden, um miteinander zu kommunizieren. Sie lernen, wie Sie HTML mit den Daten aktualisieren, die wir von APIs erhalten, mithilfe einer Technologie namens AJAX.

Die meisten Web-APIs übertragen Daten in einem Format namens JSON. JSON steht für JavaScript Object Notation.

Die JSON-Syntax sieht der JavaScript-Objektnotation sehr ähnlich. JSON hat Objekteigenschaften und deren aktuelle Werte, eingeschlossen zwischen einem `{` und einem `}`.

Diese Eigenschaften und ihre Werte werden oft als „Schlüssel-Wert-Paare“ bezeichnet.

Allerdings werden von APIs übertragene JSONs als `bytes` gesendet, und Ihre Anwendung erhält sie als `string`. Diese können in JavaScript-Objekte umgewandelt werden, sind aber standardmäßig keine JavaScript-Objekte. Die `JSON.parse`-Methode parst den String und erstellt das darin beschriebene JavaScript-Objekt.

Sie können das JSON von der freeCodeCamp Cat Photo API anfordern. Hier ist der Code, den Sie in Ihrem Klick-Ereignis verwenden können, um dies zu tun:

```js
const req = new XMLHttpRequest();
req.open('GET', '/json/cats.json', true);
req.send();
req.onload = function () {
  const json = JSON.parse(req.responseText);
  document.getElementsByClassName('message')[0].innerHTML =
    JSON.stringify(json);
};
```

Hier eine Übersicht, was jeder Teil macht. Das JavaScript-`XMLHttpRequest`-Objekt hat eine Reihe von Eigenschaften und Methoden, die zum Übertragen von Daten verwendet werden. Zuerst wird eine Instanz des `XMLHttpRequest`-Objekts erstellt und in der `req`-Variablen gespeichert. Als Nächstes initialisiert die `open`-Methode eine Anfrage – dieses Beispiel fordert Daten von einer API an, daher ist es eine `GET`-Anfrage. Das zweite Argument für `open` ist die URL der API, von der Sie Daten anfordern. Das dritte Argument ist ein Boolean-Wert, wobei `true` die Anfrage asynchron macht. Die `send`-Methode sendet die Anfrage. Schließlich parst der `onload`-Ereignishandler die zurückgegebenen Daten und wendet die `JSON.stringify`-Methode an, um das JavaScript-Objekt in einen String umzuwandeln. Dieser String wird dann als Nachrichtentext eingefügt.

# --instructions--

Aktualisieren Sie den Code, um eine `GET`-Anfrage an die freeCodeCamp Cat Photo API zu erstellen und zu senden. Klicken Sie dann auf die `Get Message`-Schaltfläche. Ihre AJAX-Funktion ersetzt den `The message will go here`-Text durch die rohe JSON-Ausgabe der API.

# --hints--

Ihr Code sollte ein neues `XMLHttpRequest` erstellen.

```js
assert.match(code, /new\s+?XMLHttpRequest\(\s*?\)/g);
```

Ihr Code sollte die `open`-Methode verwenden, um eine `GET`-Anfrage an die freeCodeCamp Cat Photo API zu initialisieren.

```js
assert.match(
  code,
  /\.open\(\s*('|")GET\1\s*,\s*('|")\/json\/cats\.json\2\s*,\s*true\s*\)/g
);
```

Ihr Code sollte die `send`-Methode verwenden, um die Anfrage zu senden.

```js
assert.match(code, /\.send\(\s*\)/g);
```

Ihr Code sollte einen `onload`-Ereignishandler haben, der auf eine Funktion gesetzt ist.

```js
assert.match(
  code,
  /\.onload\s*=\s*(function|\(\s*\))\s*(\(\s*\)|\=\>)\s*{/g
);
```

Ihr Code sollte die `JSON.parse`-Methode verwenden, um den `responseText` zu parsen.

```js
assert.match(code, /JSON\s*\.parse\(\s*.*\.responseText\s*\)/g);
```

Ihr Code sollte das Element mit der Klasse `message` holen und dessen inneres HTML in den String der JSON-Daten ändern.

```js
assert.match(
  code,
  /document\s*\.getElementsByClassName\(\s*?('|")message\1\s*?\)\[0\]\s*\.innerHTML\s*?=\s*?JSON\.stringify\(.+?\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = () => {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
