---
id: 587d7fae367417b2b2512be4
title: Auf JSON-Daten von einer API zugreifen
challengeType: 6
forumTopicId: 301499
dashedName: access-the-json-data-from-an-api
---

# --description--

In der vorherigen Aufgabe haben Sie gesehen, wie man JSON-Daten von der freeCodeCamp Cat Photo API erhält.

Jetzt schauen Sie sich die zurückgegebenen Daten genauer an, um das JSON-Format besser zu verstehen. Erinnern Sie sich an einige Notationen in JavaScript:

<blockquote>[ ] -> Eckige Klammern stellen ein Array dar.<br>{ } -> Geschweifte Klammern stellen ein Objekt dar.<br>" " -> Anführungszeichen stellen einen String dar. Sie werden auch für Schlüsselnamen in JSON verwendet.</blockquote>

Das Verständnis der Struktur der Daten, die eine API zurückgibt, ist wichtig, weil es beeinflusst, wie Sie die Werte abrufen, die Sie benötigen.

Klicken Sie rechts auf die `Get Message`-Schaltfläche, um die JSON-Daten der freeCodeCamp Cat Photo API in das HTML zu laden.

Das erste und letzte Zeichen, das Sie in den JSON-Daten sehen, sind eckige Klammern `[ ]`. Das bedeutet, dass die zurückgegebenen Daten ein Array sind. Das zweite Zeichen in den JSON-Daten ist eine geschweifte `{`-Klammer, die ein Objekt beginnt. Wenn Sie genau hinschauen, sehen Sie, dass es drei separate Objekte gibt. Die JSON-Daten sind ein Array von drei Objekten, wobei jedes Objekt Informationen über ein Katzenfoto enthält.

Sie haben zuvor gelernt, dass Objekte „Schlüssel-Wert-Paare“ enthalten, die durch Kommas getrennt sind. Im Cat Photo-Beispiel hat das erste Objekt `"id":0`, wobei `id` ein Schlüssel und `0` der entsprechende Wert ist. Ebenso gibt es Schlüssel für `imageLink`, `altText` und `codeNames`. Jedes Katzenfoto-Objekt hat diese gleichen Schlüssel, aber mit unterschiedlichen Werten.

Ein weiteres interessantes „Schlüssel-Wert-Paar“ im ersten Objekt ist `"codeNames":["Juggernaut","Mrs. Wallace","ButterCup"]`. Hier ist `codeNames` der Schlüssel und sein Wert ist ein Array von drei Strings. Es ist möglich, Arrays von Objekten zu haben, ebenso wie einen Schlüssel mit einem Array als Wert.

Erinnern Sie sich daran, wie man auf Daten in Arrays und Objekten zugreift. Arrays verwenden Klammernotation, um auf einen bestimmten Index eines Elements zuzugreifen. Objekte verwenden entweder Klammer- oder Punktnotation, um den Wert einer bestimmten Eigenschaft zuzugreifen. Hier ist ein Beispiel, das die `altText`-Eigenschaft des ersten Katzenfotos ausgibt – beachten Sie, dass die geparsten JSON-Daten im Editor in einer Variablen namens `json` gespeichert sind:

```js
console.log(json[0].altText);
```

Die Konsole würde den String `A white cat wearing a green helmet shaped melon on its head.` anzeigen.

# --instructions--

Für die Katze mit der `id` von 2 geben Sie in der Konsole den zweiten Wert im `codeNames`-Array aus. Sie sollten Klammer- und Punktnotation am Objekt verwenden (das in der Variablen `json` gespeichert ist), um auf den Wert zuzugreifen.

# --hints--

Ihr Code sollte Klammer- und Punktnotation verwenden, um den richtigen Code-Namen zuzugreifen, und `Loki` in der Konsole ausgeben.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*json\s*\[2\]\s*(\.\s*codeNames|\[\s*('|`|")codeNames\2\s*\])\s*\[\s*1\s*\]\s*\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
        // Add your code below this line

        // Add your code above this line
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
        // Add your code below this line
        console.log(json[2].codeNames[1]);
        // Add your code above this line
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
