---
id: 587d7faf367417b2b2512be9
title: Daten mit der JavaScript-XMLHttpRequest-Methode senden
challengeType: 6
forumTopicId: 301504
dashedName: post-data-with-the-javascript-xmlhttprequest-method
---

# --description--

In den vorherigen Beispielen haben Sie Daten von einer externen Ressource erhalten. Sie können auch Daten an eine externe Ressource senden, solange diese Ressource AJAX-Anfragen unterstützt und Sie die URL kennen.

Die `XMLHttpRequest`-Methode von JavaScript wird auch verwendet, um Daten an einen Server zu senden. Hier ist ein Beispiel:

```js
const xhr = new XMLHttpRequest();
xhr.open('POST', url, true);
xhr.setRequestHeader('Content-Type', 'application/json; charset=UTF-8');
xhr.onreadystatechange = function () {
  if (xhr.readyState === 4 && xhr.status === 201) {
    const serverResponse = JSON.parse(xhr.response);
    document.getElementsByClassName('message')[0].textContent =
      serverResponse.userName + serverResponse.suffix;
  }
};
const body = JSON.stringify({ userName: userName, suffix: ' loves cats!' });
xhr.send(body);
```

Sie haben einige dieser Methoden bereits gesehen. Hier initialisiert die `open`-Methode die Anfrage als `POST` an die angegebene URL der externen Ressource und übergibt `true` als dritten Parameter – was bedeutet, dass die Operation asynchron ausgeführt wird.

Die `setRequestHeader`-Methode setzt den Wert einer HTTP-Anfragekopfzeile, die Informationen über den Absender und die Anfrage enthält. Sie muss nach der `open`-Methode, aber vor der `send`-Methode aufgerufen werden. Die beiden Parameter sind der Name der Kopfzeile und der Wert, der als Inhalt dieser Kopfzeile gesetzt werden soll.

Als Nächstes behandelt der `onreadystatechange`-Ereignis-Listener eine Änderung des Zustands der Anfrage. Ein `readyState` von `4` bedeutet, dass die Operation abgeschlossen ist, und ein `status` von `201` bedeutet, dass die Anfrage erfolgreich war. Daher kann das HTML-Dokument aktualisiert werden.

Schließlich sendet die `send`-Methode die Anfrage mit dem `body`-Wert. Die `body` besteht aus einem `userName` und einem `suffix`-Schlüssel.

# --instructions--

Aktualisieren Sie den Code so, dass er eine `POST`-Anfrage an den API-Endpunkt stellt. Geben Sie dann Ihren Namen in das Eingabefeld ein und klicken Sie auf `Send Message`. Ihre AJAX-Funktion sollte `Reply from Server will be here.` durch Daten vom Server ersetzen. Formatieren Sie die Antwort so, dass Ihr Name mit dem Text ` loves cats` angehängt angezeigt wird.

# --hints--

Ihr Code sollte ein neues `XMLHttpRequest` erstellen.

```js
assert.match(code, /new\s*XMLHttpRequest\(\s*\)/g);
```

Ihr Code sollte die `open`-Methode verwenden, um eine `POST`-Anfrage an den Server zu initialisieren.

```js
assert.match(code, /\.open\(\s*('|")POST\1\s*,\s*url\s*,\s*true\s*\)/g);
```

Ihr Code sollte die `setRequestHeader`-Methode verwenden.

```js
assert.match(
  code,
  /\.setRequestHeader\(\s*('|")Content-Type\1\s*,\s*('|")application\/json;\s*charset=UTF-8\2\s*\)/g
);
```

Ihr Code sollte einen `onreadystatechange`-Ereignis-Handler haben, der auf eine Funktion gesetzt ist.

```js
assert.match(code, /\.onreadystatechange\s*=/g);
```

Ihr Code sollte das Element mit der Klasse `message` holen und dessen `textContent` auf `userName loves cats` ändern.

```js
assert.match(
  code,
  /document\.getElementsByClassName\(\s*('|")message\1\s*\)\[0\]\.textContent\s*=\s*.+?\.userName\s*\+\s*.+?\.suffix/g
);
```

Ihr Code sollte die `send`-Methode verwenden.

```js
assert.match(code, /\.send\(\s*body\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('sendMessage').onclick = function () {
      const userName = document.getElementById('name').value;
      const url = 'https://jsonplaceholder.typicode.com/posts';
      // Add your code below this line

      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Friends</h1>
<p class="message box">Reply from Server will be here</p>
<p>
  <label for="name"
    >Your name:
    <input type="text" id="name" />
  </label>
  <button id="sendMessage">Send Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('sendMessage').onclick = function () {
      const userName = document.getElementById('name').value;
      const url = 'https://jsonplaceholder.typicode.com/posts';
      // Add your code below this line
      const xhr = new XMLHttpRequest();
      xhr.open('POST', url, true);
      xhr.setRequestHeader('Content-Type', 'application/json; charset=UTF-8');
      xhr.onreadystatechange = function () {
        if (xhr.readyState === 4 && xhr.status === 201) {
          const serverResponse = JSON.parse(xhr.response);
          document.getElementsByClassName('message')[0].textContent =
            serverResponse.userName + serverResponse.suffix;
        }
      };
      const body = JSON.stringify({
        userName: userName,
        suffix: ' loves cats!'
      });
      xhr.send(body);
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Friends</h1>
<p class="message">Reply from Server will be here</p>
<p>
  <label for="name"
    >Your name:
    <input type="text" id="name" />
  </label>
  <button id="sendMessage">Send Message</button>
</p>
```
