---
id: 587d7fae367417b2b2512be5
title: JSON-Daten in HTML umwandeln
challengeType: 6
forumTopicId: 16807
dashedName: convert-json-data-to-html
---

# --description--

Da Sie jetzt Daten von einer JSON-API erhalten, können Sie diese im HTML anzeigen.

Sie können eine `forEach`-Methode verwenden, um die Daten zu durchlaufen, da die Objekte der Katzenfotos in einem Array gespeichert sind. Während Sie jedes Element durchgehen, können Sie die HTML-Elemente anpassen.

Deklarieren Sie zuerst eine html-Variable mit `let html = "";`.

Dann durchlaufen Sie das JSON und fügen der Variablen HTML hinzu, das die Schlüsselnamen in `strong`-Tags einschließt, gefolgt vom Wert. Wenn die Schleife beendet ist, rendern Sie es.

Hier ist der Code, der dies erledigt:

```js
let html = '';
json.forEach(function (val) {
  const keys = Object.keys(val);
  html += "<div class = 'cat'>";
  keys.forEach(function (key) {
    html += '<strong>' + key + '</strong>: ' + val[key] + '<br>';
  });
  html += '</div><br>';
});
```

**Hinweis:** Für diese Aufgabe müssen Sie neue HTML-Elemente auf der Seite hinzufügen, daher können Sie sich nicht auf `textContent` verlassen. Stattdessen müssen Sie `innerHTML` verwenden, was eine Website anfällig für Cross-Site-Scripting-Angriffe machen kann.

# --instructions--

Fügen Sie eine `forEach`-Methode hinzu, um die JSON-Daten zu durchlaufen und die HTML-Elemente zu erstellen, um sie anzuzeigen.

Hier ist ein Beispiel-JSON:

```json
[
  {
    "id": 0,
    "imageLink": "https://cdn.freecodecamp.org/curriculum/legacy-json-apis-ajax/funny-cat.jpg",
    "altText": "A white cat wearing a green helmet shaped melon on its head. ",
    "codeNames": ["Juggernaut", "Mrs. Wallace", "Buttercup"]
  }
]
```

# --hints--

Ihr Code sollte die Daten in der `html`-Variablen speichern

```js
assert.match(__helpers.removeWhiteSpace(code), /html(\+=|=html\+)/g);
```

Ihr Code sollte eine `forEach`-Methode verwenden, um die JSON-Daten von der API zu durchlaufen.

```js
assert.match(code, /json\.forEach/g);
```

Ihr Code sollte die Schlüsselnamen in `strong`-Tags einschließen.

```js
assert.match(code, /<strong>.+<\/strong>/g);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        let html = '';
        // Add your code below this line

        // Add your code above this line
        document.getElementsByClassName('message')[0].innerHTML = html;
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        let html = '';
        // Add your code below this line
        json.forEach(function (val) {
          var keys = Object.keys(val);
          html += "<div class = 'cat'>";
          keys.forEach(function (key) {
            html += '<strong>' + key + '</strong>: ' + val[key] + '<br>';
          });
          html += '</div><br>';
        });
        // Add your code above this line
        document.getElementsByClassName('message')[0].innerHTML = html;
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
