---
id: 5ccfad82bb2dc6c965a848e5
title: JSON mit der JavaScript-fetch-Methode abrufen
challengeType: 6
forumTopicId: 301501
dashedName: get-json-with-the-javascript-fetch-method
---

# --description--

Eine weitere Möglichkeit, externe Daten anzufordern, ist die Verwendung der JavaScript `fetch()`-Methode. Sie entspricht `XMLHttpRequest`, aber die Syntax gilt als leichter verständlich.

Hier ist der Code, um eine GET-Anfrage an `/json/cats.json` zu stellen

```js
fetch('/json/cats.json')
  .then(response => response.json())
  .then(data => {
    document.getElementById('message').innerHTML = JSON.stringify(data);
  });
```

Hinweis: Die `fetch()`-Methode verwendet standardmäßig `GET` als `HTTP`-Methode. Das bedeutet, dass Sie diese für einfache Datenabfragen nicht explizit angeben müssen.

Sehen wir uns jeden Teil dieses Codes an.

Die erste Zeile ist die, die die Anfrage stellt. Also macht `fetch(URL)` eine `GET`-Anfrage an die angegebene URL. Die Methode gibt ein Promise zurück.

Nachdem ein Promise zurückgegeben wurde, wird, falls die Anfrage erfolgreich war, die `then`-Methode ausgeführt, die die Antwort nimmt und in das JSON-Format umwandelt.

Die `then`-Methode gibt ebenfalls ein Promise zurück, das von der nächsten `then`-Methode verarbeitet wird. Das Argument im zweiten `then` ist das JSON-Objekt, das Sie suchen!

Jetzt wird das Element ausgewählt, das die Daten erhalten soll, indem `document.getElementById()` verwendet wird. Dann wird der HTML-Code des Elements geändert, indem ein String eingefügt wird, der aus dem JSON-Objekt erstellt wurde, das von der Anfrage zurückgegeben wurde.

# --instructions--

Aktualisieren Sie den Code, um eine `GET`-Anfrage an die freeCodeCamp Cat Photo API zu erstellen und zu senden. Aber diesmal verwenden Sie die `fetch`-Methode anstelle von `XMLHttpRequest`.

# --hints--

Ihr Code sollte die abgerufenen Daten verwenden, um das innere HTML zu ersetzen

```js
const catData = 'dummy data';
const ref = fetch;
fetch = () => Promise.resolve({ json: () => catData });

  try {
    document.getElementById('getMessage').click();
    await new Promise((resolve, reject) => setTimeout(() => resolve(), 250));
  } catch (error) {
    console.log(error);
  } finally {
    fetch = ref;
    assert.equal(
      document.getElementById('message').textContent,
      JSON.stringify(catData)
    );
  }

```

Ihr Code sollte eine `GET`-Anfrage mit `fetch` stellen.

```js
assert.match(code, /fetch\s*\(\s*('|")\/json\/cats\.json\1\s*\)/g);
```

Ihr Code sollte `then` verwenden, um die Antwort in JSON umzuwandeln.

```js
assert.match(
  code,
  /\.then\s*\(\s*\(?(?<var>\w+)\)?\s*=>\s*\k<var>\s*\.json\s*\(\s*\)\s*\)/g
);
```

Ihr Code sollte `then` verwenden, um die in JSON umgewandelten Daten mit der anderen `then` zu verarbeiten.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.then\(\(?\w+\)?=>{[^}]*}\)/g);
```

Ihr Code sollte das Element mit der ID `message` holen und dessen inneres HTML in den String der JSON-Daten ändern.

```js
assert.match(
  __helpers.removeWhiteSpace(code),
  /document\.getElementById\(('|")message\1\)\.innerHTML=JSON\.stringify\(?\w+\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = () => {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>
<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p id="message" class="box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = () => {
      fetch('/json/cats.json')
        .then(response => response.json())
        .then(data => {
          document.getElementById('message').innerHTML = JSON.stringify(data);
        });
    };
  });
</script>
<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p id="message" class="box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
