---
id: 687d7deea6a2a22fa8816a73
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

Nelle lezioni precedenti, hai imparato alcune strutture dati comuni. In questo workshop, implementerai una lista concatenata.

Per prima cosa, crea una classe chiamata `LinkedList` per rappresentare la struttura che conterrà i nodi della tua lista. Aggiungi la parola chiave `pass` all'interno della classe in modo che il test possa passare.

# --hints--

Dovresti creare una classe chiamata `LinkedList`. Ricorda di aggiungere la parola chiave `pass`.

```js
({
    test: () => runPython(`assert _Node(_code).has_class("LinkedList")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
