---
id: 687d8e05a529a4838f1c961c
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Successivamente, rimuovi la parola chiave `pass` all'interno della classe `Node`, poi dai alla classe `Node` un metodo `__init__` con due parametri, `self` e `element`.

All'interno del metodo, assegna `element` a `self.element`. Questo archivia il valore passato durante la creazione di un nodo.

# --hints--

Dovresti definire un metodo `__init__` all'interno della tua classe `Node`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").has_function("__init__")`) 
})
```

Il tuo metodo `__init__` dovrebbe avere due parametri, `self` e `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").find_function("__init__").has_args("self, element")`) 
})
```

Dovresti assegnare `element` a `self.element` all'interno del metodo `__init__`.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.element == "A"
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
--fcc-editable-region--
    class Node:
        pass
--fcc-editable-region--
```
