---
id: 687d8c6f066877768f7ee441
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

All'interno della tua classe `LinkedList`, sostituisci la parola chiave `pass` con una classe interna chiamata `Node`. Questa rappresenterà i singoli nodi che archiviano dati nella lista.

# --hints--

Dovresti creare una classe `Node` all'interno della classe `LinkedList`. Ricorda di aggiungere la parola chiave `pass` dentro la classe `Node`.

```js
({
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").has_class("Node")`) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class LinkedList:
    pass
--fcc-editable-region--
```
