---
id: 687e366b3004a9831e4cbd41
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

Aggiungi un metodo chiamato `is_empty` che restituisce `self.length == 0`.

Lo userai per tenere traccia di quanti nodi ci sono nella lista. Quindi, con `self.length == 0`, restituisce `True`, il che significa che la lista non ha nodi.

Se `self.length` è maggiore di `0`, restituisce `False`, il che significa che la lista ha almeno un nodo.

# --hints--

Dovresti creare un metodo `is_empty` che restituisce `self.length == 0`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.is_empty() == True
    ll.length = 1
    assert ll.is_empty() == False`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

--fcc-editable-region--
    
--fcc-editable-region--
```
