---
id: 687e344cc1993f7316af030c
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

All'interno della classe `LinkedList`, aggiungi un metodo `__init__` che imposta `self.length` su `0`. Questo inizializza un contatore per tenere traccia di quanti nodi (elementi) sono attualmente nella lista.

# --hints--

Dovresti impostare `self.length` su `0` all'interno del metodo `__init__` della classe `LinkedList`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.length == 0`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    
--fcc-editable-region--
```
