---
id: 688003f589b875af60273dd9
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Sempre all'interno del metodo `__init__`, assegna `None` a `self.next`. Questo imposta il puntatore `next` su `None`, con `None` che indica che il nuovo nodo non punta ancora a nessun altro nodo nella lista.

Successivamente, `next` può essere aggiornato per puntare a un altro nodo e formare la catena di una lista concatenata.

# --hints--

All'interno del metodo `__init__`, dovresti inizializzare l'attributo `next` a `None`.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.next is None
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
--fcc-editable-region--
        def __init__(self, element):
            self.element = element
--fcc-editable-region--
```
