---
id: 68800a64b06e80d9a31c131b
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Sempre nel metodo `__init__` della classe `LinkedList`, imposta `self.head` su `None`. Questo agisce come riferimento al primo nodo della lista. La lista ├© vuota al momento quindi devi impostarlo su `None`.


# --hints--

Dovresti impostare `self.head` su `None` all'interno del metodo `__init__` della classe `LinkedList`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.head is None`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    def __init__(self):
        self.length = 0
--fcc-editable-region--
```
