---
id: 68857e0a49baf4751f737a2a
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

Ora, dopo il metodo `is_empty`, crea un metodo chiamato `add` all'interno della classe `LinkedList`. Il nuovo metodo dovrebbe avere due parametri, `self` e `element`.

All'interno del nuovo metodo, crea una variabile chiamata `node` e assegnale una nuova istanza di `Node` usando `self.Node(element)`. Questo creerà un nuovo nodo da aggiungere alla lista concatenata.

# --hints--

Dovresti creare un metodo `add` all'interno della classe `LinkedList`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").has_function("add")`
      )
    )
})
```

Il tuo metodo `add` dovrebbe avere due parametri, `self` e `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("add").has_args("self, element")`) 
})
```

Dovresti assegnare `self.Node(element)` a `node` nel metodo `add`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_body().is_equivalent("node = self.Node(element)")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
