---
id: 688649c8673f842f49228a5e
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Se la lista non è vuota, il metodo dovrebbe trovare l'ultimo nodo e collegare il nuovo nodo al suo puntatore `next`.

Per farlo, crea un blocco `else` dopo il blocco `if`. All'interno del blocco `else`, crea una variabile `current_node` e assegnale `self.head`.

# --hints--

Dovresti assegnare `self.head` a `current_node` all'interno del blocco `else`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].is_equivalent("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
