---
id: 68867c3898b09115d8c5140f
title: Passaggio 13
challengeType: 20
dashedName: step-13
---

# --description--

Il loop `while` che hai appena creato ti permette di iterare su ogni nodo della lista fino a raggiungere l'ultimo nodo. Una volta fatto ciò, puoi aggiungere il nodo alla lista.

Nel blocco `else`, subito dopo il blocco `while`, assegna `node` a `current_node.next`.

# --hints--

Dovresti assegnare `node` a `current_node.next` nel blocco `else`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].has_stmt("current_node.next = node")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
