---
id: 6888d05ba074461b64f8eaca
title: Passaggio 16
challengeType: 20
dashedName: step-16
---

# --description--

Ora creerai un altro metodo che servirà per rimuovere un elemento specifico dalla lista concatenata.

Crea un metodo chiamato `remove` con due parametri, `self` e `element`. Al suo interno, crea una variabile chiamata `previous_node` e assegnale `None`.

# --hints--

Dovresti creare un metodo chiamato `remove`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_class("LinkedList").has_function("remove")
    `)
  )
})
```

Il tuo metodo `remove` dovrebbe avere due parametri, `self` e `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("remove").has_args("self, element")`) 
})
```

Dovresti assegnare `None` a `previous_node` nel metodo `remove`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("remove").find_body().is_equivalent("previous_node = None")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
