---
id: 68b84df2648b4c96b25022f3
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

Mettiamo alla prova la tua lista concatenata finora! In fondo al tuo codice, crea un'istanza di `LinkedList` chiamata `my_list`. Poi, usando il metodo `is_empty` che hai creato nell'ultimo passaggio, stampa `my_list.is_empty()`.

# --hints--

Dovresti creare un'istanza di `LinkedList` chiamata `my_list`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_variable("my_list").value_is_call("LinkedList")
    `)
  )
})
```

Dovresti stampare `my_list.is_empty()`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).has_call("print(my_list.is_empty())")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0

--fcc-editable-region--

--fcc-editable-region--
```
