---
id: 68b9e951e4892721632acc26
title: Passaggio 15
challengeType: 20
dashedName: step-15
---

# --description--

Ottimo! Ora mettiamo alla prova l’aggiunta di elementi alla tua lista concatenata. In fondo al tuo codice, dopo aver creato un’istanza di `LinkedList` chiamata `my_list` e aver stampato `my_list.is_empty()`, aggiungi `1` e `2` alla lista usando il metodo `add`.

Poi, stampa di nuovo `my_list.is_empty()` per confermare che la lista non è più vuota. Infine, stampa `my_list.length` per vedere la lunghezza attuale della lista concatenata.

# --hints--

Dovresti aggiungere `1` alla lista usando il metodo `add`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(1)")`)
  )
})
```

Dovresti aggiungere `2` alla lista usando il metodo `add`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(2)")`)
  )
})
```

Dovresti stampare `my_list.is_empty()`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[1].is_equivalent("print(my_list.is_empty())")
    `)
  )
})
```

Dovresti stampare `my_list.length`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[2].is_equivalent("print(my_list.length)")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

my_list = LinkedList()
print(my_list.is_empty())

--fcc-editable-region--

--fcc-editable-region--
```
