---
id: 68ba19dec78f654b7fca401e
title: Passaggio 24
challengeType: 20
dashedName: step-24
---

# --description--

Ora, testiamo la rimozione di un elemento dalla tua lista concatenata. Alla fine del tuo codice, rimuovi `1` dalla lista usando il metodo `remove`. Poi, stampa di nuovo `my_list.length` per vedere la lunghezza aggiornata della lista concatenata.

Con questo, il workshop sulla struttura dati lista concatenata è completo!

# --hints--

Dovresti rimuovere `1` dalla lista usando il metodo `remove`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.remove(1)")`)
  )
})
```

Dovresti stampare `my_list.length`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[3].is_equivalent("print(my_list.length)")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
        while current_node is not None and current_node.element != element:
            previous_node = current_node
            current_node = current_node.next
        if current_node is None:
            return        
        elif previous_node is not None:
            previous_node.next = current_node.next
        else:
            self.head = current_node.next
        self.length -= 1

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0

    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
        while current_node is not None and current_node.element != element:
            previous_node = current_node
            current_node = current_node.next
        if current_node is None:
            return
        elif previous_node is not None:
            previous_node.next = current_node.next
        else:
            self.head = current_node.next
        self.length -= 1

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)

my_list.remove(1)
print(my_list.length)
```
