---
id: 68c0da2df8ce865dcf7769a7
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

Ancora all'interno del blocco `else`, crea un loop `while` che verifichi che `current_node.next` non sia `None`.

All'interno del loop `while`, assegna `current_node.next` a `current_node`.

# --hints--

Dovresti creare un loop `while` che controlli se `current_node.next` non è `None`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].find_whiles()[0].find_conditions()[0]
assert _cond.is_equivalent("current_node.next is not None")
`) })
```

Dovresti assegnare `current_node.next` a `current_node` all'interno del loop `while`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].find_whiles()[0]
assert _cond.is_equivalent("while current_node.next is not None:\\n current_node = current_node.next")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--
        else:
            current_node = self.head
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
