---
id: 7a97b79c096ecb000d03bd36
title: Einen Rechner für die n-te Fibonacci-Zahl erstellen
challengeType: 26
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `fibonacci` erstellen.
2. Sie sollten innerhalb der Funktion `sequence` ein Array namens `fibonacci` definieren, das mit den Werten `[0, 1]` initialisiert ist.
3. Die Funktion `fibonacci` sollte einen Parameter akzeptieren, eine nicht-negative ganze Zahl `n`.
4. Das Aufrufen von `fibonacci(n)` sollte einen dynamischen Programmieransatz verwenden, um die `n`-te Zahl der Fibonacci-Folge zu berechnen und zurückzugeben, wobei jede Zahl die Summe der beiden vorhergehenden Zahlen ist.
5. Jede berechnete Fibonacci-Zahl sollte an das Array `sequence` angehängt werden.

# --hints--

Sie sollten eine Funktion mit dem Namen `fibonacci` haben.

```js
assert.isFunction(fibonacci);
```

Ihre Funktion `fibonacci` sollte einen einzelnen Parameter annehmen.

```js
const params = __helpers.getFunctionParams(fibonacci.toString());
assert.equal(params.length, 1);
```

Sie sollten innerhalb der Funktion `sequence` ein Array namens `fibonacci` haben, das mit `[0, 1]` initialisiert ist.

```js
assert.match(__helpers.removeJSComments(fibonacci.toString()), /(const|let|var)\s+sequence\s*=\s*\[\s*0\s*,\s*1\s*\]/);
```

`fibonacci(0)` sollte `0` zurückgeben.

```js
assert.equal(fibonacci(0), 0);
```

`fibonacci(1)` sollte `1` zurückgeben.

```js
assert.equal(fibonacci(1), 1);
```

`fibonacci(2)` sollte `1` zurückgeben.

```js
assert.equal(fibonacci(2), 1);
```

`fibonacci(3)` sollte `2` zurückgeben.

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` sollte `5` zurückgeben.

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` sollte `55` zurückgeben.

```js
assert.equal(fibonacci(10), 55);
```

`fibonacci(15)` sollte `610` zurückgeben.

```js
assert.equal(fibonacci(15), 610);
```

Sie sollten in Ihrem Code keine Rekursion verwenden.

```js
const funcStr = __helpers.removeJSComments(fibonacci.toString());
const bodyMatch = funcStr.match(/^[^{]*\{([\s\S]*)\}[^}]*$/) || funcStr.match(/=>\s*([\s\S]+)$/);
assert(bodyMatch && !bodyMatch[1].match(/\bfibonacci\s*\(/));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fibonacci(n) {
  const sequence = [0, 1];

  if (n < 2) {
    return sequence[n];
  }

  for (let i = 2; i <= n; i++) {
    sequence.push(sequence[i - 1] + sequence[i - 2]);
  }

  return sequence[n];
}
```
