---
id: 684cd027744c2a0520ba4d58
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

Tailwind CSS は、マークアップに直接適用されるユーティリティクラスを使います。これらのクラスには、フォント、色、Flexbox レイアウト、Grid などのスタイリングが含まれています。

Tailwind CSS で背景色を適用するには、`bg` 接頭辞の後に色の名前と数字を続けて使います。この数字は色の濃淡を表します。

以下は、段落要素に濃い青の背景色を適用する例です。

```html
<p class="bg-blue-900">This has a dark blue background.</p>
```

開く `div` タグに、薄い赤の背景色を表すユーティリティクラス `bg-red-100` を含むクラスを追加してください。

# --hints--

開く `div` タグには class 属性が必要です。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.hasAttribute("class"));
```

class 属性にはユーティリティクラス `bg-red-100` を含めてください。

```js
const divEl = document.querySelector("div");
assert.exists(document.querySelector("div.bg-red-100"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
