---
id: 684cc59eed2706fa50e8933c
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、エラーメッセージのコンポーネントを作成しながら、Tailwind CSSの基本的な使い方を学びます。

最初の数ステップでは、初期のHTMLマークアップの作成にフォーカスします。

`body` 要素の中に `div` 要素を追加してください。次に、その `div` 要素の中に `p` 要素を追加します。

# --hints--

`div` 要素があるはずです。

```js
assert.exists(document.querySelector("div"));
```

`p` 要素があるはずです。

```js
assert.exists(document.querySelector("p"));
```

`p` 要素は `div` 要素の中にあるべきです。

```js
assert.exists(document.querySelector("div p"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
