---
id: 684cd59ec94b5e063bf776f0
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

薄い赤の背景の周りにボーダーがあるといいですね。

Tailwind CSSで`border-width`を設定するには、`border-number`ユーティリティクラスを使えます。ここで`number`はボーダーの明度や太さを表します。

すべての辺に`border-width`の`4px`を適用する例を示します。

```html
<p class="border-4">This is how to set the border-width.</p>
```

開く`div`タグの中に、すべての辺に`border-width`の`2px`を設定する正しいユーティリティクラスを追加してください。

# --hints--

開く`div`タグには`border-2`クラスが必要です。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
