---
id: 684cd8bf5445de089af9c837
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

次はボーダーの色を設定する時です。Tailwind CSSで`border-color`プロパティを設定するには、次のパターンのユーティリティクラスを使います：`border-color-number`。ここで`number`はボーダー色の色合いを表します。

要素に紫のボーダー色を適用する例を示します：

```html
<p class="border-purple-400">This is border-color.</p>
```

開く`div`タグの中で、赤のボーダー色を`300`の色合いで適用してください。

# --hints--

開く`div`タグには`border-red-300`ユーティリティクラスが必要です。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-red-300"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
