---
id: 684cca09aac0fc035db4099f
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`p` 要素の中に、`Error! Something went wrong. Please try again.` が `Error!` 要素で囲まれている場所に `strong` というテキストを追加してください。

# --hints--

`strong` 要素の中に `p` 要素があるはずです。

```js
assert.exists(document.querySelector("p strong"));
```

`strong` 要素には `Error!` というテキストが含まれているはずです。

```js
assert.equal(document.querySelector("strong")?.textContent, "Error!");
```

`p` 要素には `Something went wrong. Please try again.` で終わる正しいテキストが含まれているはずです。

```js
const remainingText = "Something went wrong. Please try again."
const pElement = document.querySelector("p");
assert.isTrue(pElement?.textContent.trim().endsWith(remainingText));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      --fcc-editable-region--
      <p>

      </p>
      --fcc-editable-region--
    </div>
  </body>
</html>
```
