---
id: 684cff0bf0a4ca0ad7e42990
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

次のステップはボーダーの角を丸くすることです。Tailwind CSSで`border-radius`プロパティを設定するには、`rounded-size`が半径を表す次のパターンのユーティリティクラスを使うことができます：`size`。

要素に特大のボーダー半径を適用する例を示します：

```html
<p class="rounded-xl">This is border radius.</p>
```

開く`div`タグの中で、`md`のサイズのボーダー半径を適用してください。

# --hints--

開く`div`タグには`rounded-md`ユーティリティクラスを付ける必要があります。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("rounded-md"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
