---
id: 684ccbdc1ef89e042b744e09
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

ワークショップの次の部分は、閉じるボタンを追加することです。

`p` 要素の下に、`button` が `aria-label` に設定された `"Dismiss"` 要素を追加してください。その `button` 要素の中に、`span` 属性を持つ `aria-hidden` 要素を追加します。`span` 要素のテキストは `&times;` であるべきです。

*注意*: このワークショップは Tailwind CSS のスタイル操作にのみフォーカスするため、`button` にはインタラクティブな機能は追加されません。

# --hints--

`button` 要素があるはずです。

```js
assert.exists(document.querySelector("button"));
```
 
`button` 要素には `aria-label` 属性があるべきです。

```js
const buttonEl = document.querySelector("button");
assert.isTrue(buttonEl?.hasAttribute("aria-label"));
```

`button` 要素には `aria-label` 属性が `"Dismiss"` に設定されているべきです。

```js
const buttonEl = document.querySelector("button");
assert.equal(buttonEl?.getAttribute("aria-label"), "Dismiss");
```

`button` 要素の中に `span` 要素がネストされているべきです。

```js
assert.exists(document.querySelector("button span"));
```

`span` 要素には `aria-hidden` 属性があるべきです。

```js
const spanEl = document.querySelector("span");
assert.isTrue(spanEl?.hasAttribute("aria-hidden"));
```

`span` 要素のテキストは `&times;` であるべきです。

```js
const spanEl = document.querySelector("span");
assert.equal(spanEl?.innerHTML, "×");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--

      --fcc-editable-region--
    </div>
  </body>
</html>
```
