---
id: 684d160a2370b01310ae0fda
title: ステップ 18
challengeType: 0
dashedName: step-18
---

# --description--

新しい赤色はきれいですが、テキストがもっと大きいほうが良いです。Tailwind CSSを使って要素の`font-size`を設定するには、次のパターンを使えます：`text-value`。

ここに、要素の`font-size`を`0.875rem`または（`14px`）に設定する例があります：

```html
<p class="text-sm">This is for font-size.</p>
```

開く`p`要素の中で、`font-size`を`xl`に設定する正しいユーティリティクラスを使ってください。これは`1.25rem`または（`20px`）を表します。

# --hints--

開く`p`タグには`text-xl`ユーティリティクラスを付けてください。

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl.classList.contains("text-xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p class="text-red-700">
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
