---
id: 684d0d369e2b3b0f288d2c20
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

`div` 要素を水平方向に中央揃えにすると良いでしょう。`margin` を両側に同時に適用するには、次のパターンを使えます：`mx-number` または `my-number`。

`mx` は左右のマージンを設定し、`my` は上下のマージンを設定します。

以下は、要素に上下のマージンを適用する例です：

```html
<p class="my-8">This is margin.</p>
```

開く `div` タグの中で、左右のマージンを `auto` に設定する正しいユーティリティクラスを使ってください。

# --hints--

開く `div` タグには `mx-auto` ユーティリティクラスを付けるべきです。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
