---
id: 684d105e488dbd1166a717ae
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

プレビューウィンドウのサイズを大きくすると、すべてのアイテムがコンテナ内で左揃えのままであることに気づきます。フレックスアイテムが中央に配置されていると良いですね。

Tailwind CSSでコンテナの主軸に沿ってフレックスアイテムを配置するには、次のパターンを使います：`justify-value`。`value`は`start`、`end`、`center`、`between`などが使えます。`justify-content`でバニラCSSで使える有効な値は、Tailwind CSSでもここで使えます。

開く`div`タグの中に、`justify-content`プロパティを`center`に設定する正しいユーティリティクラスを使ってください。

# --hints--

開く`div`タグには`justify-center`ユーティリティクラスを付けてください。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("justify-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
