---
id: 684d17a3b350bb13d3c19c0f
title: ステップ 19
challengeType: 0
dashedName: step-19
---

# --description--

`button` 要素にスタイルを適用する時間です。

まず、`button` 要素にクラス属性を追加してください。次に、赤色のテキストカラーを設定するために、色の濃さを `700` に設定する正しいユーティリティクラスを適用してください。

それから、開く `button` タグの中で、`font-size` を `2xl` に設定する正しいユーティリティクラスを使ってください。これは `1.5rem` または (`24px`) を表します。

# --hints--

`button` 要素にはクラス属性が必要です。

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.hasAttribute("class"));
```

開く `button` タグには `text-red-700` ユーティリティクラスを付けてください。

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.classList.contains("text-red-700"));
```

開く `button` タグには `text-2xl` ユーティリティクラスを付けてください。

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl.classList.contains("text-2xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--
      <button aria-label="Dismiss">
      --fcc-editable-region--
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
