---
id: 684d07cbfa43a20d6a002ed7
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

`div` 要素のスタイルがだんだん整ってきました。でも、テキストの周りに少し余白があるといいですね。

Tailwind CSSで要素に`padding`を適用するには、次のパターンを使います：`p-number`。ここで`number`は、4辺すべてに適用したいピクセル数を表します。

要素に`padding`の`8px`を適用する例を示します：

```html
<p class="p-8">This is padding.</p>
```

開く`div`タグの中で、正しいユーティリティクラスを使って`padding`の`4px`を適用してください。

# --hints--

開く`div`タグには`p-4`ユーティリティクラスを付けてください。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
