---
id: 684d009994a3840b98c9aa3a
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

現在、`div` 要素の幅はページ全体の幅に広がっています。Tailwind CSSで幅を指定するには、固定幅に`w-number` パターンを使ったり、パーセンテージに`w-fraction` パターンを使ったりする方法があります。

親コンテナの`66%` に幅を設定する例を示します。

```html
<p class="w-2/3">This is a percentage width.</p>
```

開く`div` タグの中に、親コンテナの`50%` に幅を設定してください。

# --hints--

開く`div` タグには`w-1/2` ユーティリティクラスを付けてください。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
