---
id: 684d1942a1904a14a3709c6c
title: ステップ 20
challengeType: 0
dashedName: step-20
---

# --description--

エラーメッセージコンポーネントのスタイリングはほぼ完成していますが、修正すべき点が1つだけ残っています。段落とボタンのテキストで異なるフォントサイズを設定したため、これらのフレックスアイテム間の縦方向の整列がずれてしまっています。

Tailwind CSSでコンテナの交差軸に沿ってフレックスアイテムを配置するには、次のパターンを使うことができます：`items-value`。`value`は`start`、`end`、`center`、`baseline`など、さまざまなものが使えます。

開く`div`タグの中で、`align-items`プロパティを`center`に設定する正しいユーティリティクラスを使ってください。

この最後の変更で、エラーメッセージコンポーネントは完成です！

# --hints--

開く`div`タグには`items-center`ユーティリティクラスを付けてください。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
    --fcc-editable-region--
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto flex gap-4 justify-center items-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
