---
id: 684d03df2a0ddf0c735865f0
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

今、要素の幅が`body`の半分になっているのがわかります。これは大きなデバイスではうまく機能しますが、小さなデバイスではあまりうまくいきません。

Tailwind CSSはモバイルファーストの設計方針を採用しています。したがって、他のデバイス向けにレスポンシブデザインを作成したい場合は、`md`や`lg`のようなTailwindのレスポンシブバリアントを次のように使うことができます。

```html
<p class="lg:w-2/3">This is a percentage width.</p>
```

次の例では、デバイスサイズが`66%`以上の場合に親要素の幅を`1024px`に設定します。

開く`div`タグの中で、`w-1/2`を修正して、デバイスサイズが`768px`以上の場合にのみ適用されるようにしてください。このクラスには`md`バリアントを接頭辞として付けることができます。

# --hints--

開く`div`タグには`md:w-1/2`ユーティリティクラスを付けてください。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:w-1/2"));
```

開く`w-1/2`タグにはもう`div`クラスが含まれていないはずです。

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
