---
id: 684d11ffd15a531225d081e9
title: ステップ 17
challengeType: 0
dashedName: step-17
---

# --description--

`div` 要素にスタイルを適用したので、次は `p` 要素にスタイルを適用しましょう。

現在、段落要素のテキスト色は黒ですが、暗い赤色にすると良いでしょう。

前のステップで、背景色の扱い方を学びました：`bg-red-100`。テキスト色も同じように、次のパターンに従って設定できます：`text-color-shade`。

以下は、要素に青いテキスト色を適用する例です：

```html
<p class="text-blue-500">This has blue text.</p>
```

まず、`p` 要素にクラス属性を追加してください。その後、色の濃さを `700` に設定した赤いテキスト色用のユーティリティクラスを適用しましょう。

# --hints--

`p` 要素にはクラス属性が必要です。

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.hasAttribute("class"));
```

開始タグの `p` には `text-red-700` ユーティリティクラスを付けてください。

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.classList.contains("text-red-700"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p>
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
