---
id: 684d09cad4ec900e41db2e74
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

現在、`div` 要素はプレビュー画面の左上隅に配置されています。しかし、要素の上部に少しスペースを追加すると良いでしょう。Tailwind CSSで要素に`margin`を適用するには、次のパターンを使えます：`m-number`。ここで`number`は、4辺すべてに適用する固定ピクセル数を表します。

もしマージンを片側だけに適用したい場合は、次のいずれかのパターンを使えます：`mt-number`（上マージン）、`mb-number`（下マージン）、`mr-number`（右マージン）、`ml-number`（左マージン）。

要素に`margin-bottom`を適用する例を示します：

```html
<p class="mb-8">This is margin.</p>
```

開く`div`タグの中で、`margin-top`の`4px`を適用する正しいユーティリティクラスを使ってください。

# --hints--

開く`div`タグには`mt-4`ユーティリティクラスを付けるべきです。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mt-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
