---
id: 684d0e11094c3e0fe1fb98f5
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

現在、段落要素は dismiss `button` 要素の上に積み重なっています。しかし、`button` 要素を段落と同じ行に並べると見栄えが良くなります。ここで Flexbox が役立ちます。

Tailwind CSS で `div` 要素を Flexbox レイアウトに設定するには、`flex` ユーティリティクラスを使います。

# --hints--

開く `div` タグには `flex` ユーティリティクラスを付けてください。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("flex"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
