---
id: 684d0eeb9c81e5109e52fe8b
title: ステップ 15
challengeType: 0
dashedName: step-15
---

# --description--

`div` 要素は見た目が良くなってきました。しかし、dismiss `button` が段落要素に少し近すぎます。

フレックスアイテム間にスペースを作るには、次のパターンを使うことができます：`gap-number`。`number` はフレックスレイアウトの行と列の間の固定スペースを表します。

要素に `gap` を適用する例を示します：

```html
<p class="flex gap-8">This is gap.</p>
```

開く `div` タグの中で、`gap` を `4` に設定する正しいユーティリティクラスを使ってください。

# --hints--

開く `div` タグには `gap-4` ユーティリティクラスを付ける必要があります。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("gap-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
