---
id: 68501d32b69f9d7756334ffc
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

プレビューウィンドウを拡大すると、`div` 要素がページの中央に水平配置されているのが見えるはずです。しかし、小さいデバイスでは、`div` 要素の幅がページ全体の幅に広がっています。

`mx-auto` クラスを、`md` デバイス以上にのみ適用されるように更新してください。

# --hints--

開く `div` タグには `md:mx-auto` ユーティリティクラスを付けてください。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:mx-auto"));
```

開く `mx-auto` タグに `div` クラスがもう含まれていないことを確認してください。

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
