---
id: 587d7db7367417b2b2512b9d
title: 匹配開頭字串樣式
challengeType: 1
forumTopicId: 301349
dashedName: match-beginning-string-patterns
---

# --description--

先前的挑戰顯示，regular expressions 可以用來尋找多個符合項目。它們也用於在字串中特定位置搜尋樣式。

在先前的挑戰中，你在字元集合內使用了插入符號字元（`^`）來創建一個否定字元集合，形式為 `[^thingsThatWillNotBeMatched]`。在字元集合外，插入符號用於搜尋字串開頭的樣式。

```js
let firstString = "Ricky is first and can be found.";
let firstRegex = /^Ricky/;
firstRegex.test(firstString);
let notFirst = "You can't find Ricky now.";
firstRegex.test(notFirst);
```

第一次 `test` 呼叫會傳回 `true`，而第二次會傳回 `false`。

# --instructions--

在正規表示式中使用插入字元來尋找字串 `rickyAndCal` 開頭的 `Cal`。

# --hints--

你的正規表達式應該搜尋以大寫字母開頭的字串 `Cal`。

```js
assert(calRegex.source == '^Cal');
```

你的正規表達式不應該使用任何旗標。

```js
assert(calRegex.flags == '');
```

你的正則表達式應該在字串開頭匹配 `Cal`。

```js
calRegex.lastIndex = 0;
assert(calRegex.test('Cal and Ricky both like racing.'));
```

你的正則表達式不應該匹配字串中間的 `Cal`。

```js
calRegex.lastIndex = 0;
assert(!calRegex.test('Ricky and Cal both like racing.'));
```

# --seed--

## --seed-contents--

```js
let rickyAndCal = "Cal and Ricky both like racing.";
let calRegex = /change/; // Change this line
let result = calRegex.test(rickyAndCal);
```

# --solutions--

```js
let rickyAndCal = "Cal and Ricky both like racing.";
let calRegex = /^Cal/; // Change this line
let result = calRegex.test(rickyAndCal);
```
