---
id: 587d7db7367417b2b2512b9c
title: 在追捕中找到一個或多個罪犯
challengeType: 1
forumTopicId: 301343
dashedName: find-one-or-more-criminals-in-a-hunt
---

# --description--

是時候暫停並測試你新的正則表達式撰寫技巧。一群罪犯從監獄逃脫並逃跑了，但你不知道有多少人。然而，你知道他們在其他人附近時會保持靠近。你負責一次找到所有罪犯。

這裡有一個範例來檢視如何執行這個操作：

正則表達式 `/z+/` 會匹配連續出現一次或多次的字母 `z`。它會在以下所有字串中找到匹配項：

```js
"z"
"zzzzzz"
"ABCzzzz"
"zzzzABC"
"abczzzzzzzzzzzzzzzzzzzzzabc"
```

但它不會在以下字串中找到符合項，因為沒有字母 `z` 字元：

```js
""
"ABC"
"abcabc"
```

# --instructions--

撰寫一個貪婪的正則表達式，以尋找群組中一個或多個罪犯。罪犯由大寫字母 `C` 表述。

# --hints--

你的正規表示式應該在字串 `C` 中匹配一個罪犯（`C`）。

```js
assert('C'.match(reCriminals) && 'C'.match(reCriminals)[0] == 'C');
```

你的正規表達式應該在字串 `CC` 中匹配兩個罪犯（`CC`）。

```js
assert('CC'.match(reCriminals) && 'CC'.match(reCriminals)[0] == 'CC');
```

你的正規表達式應該在字串 `P1P5P4CCCcP2P6P3` 中匹配三個罪犯（`CCC`）。

```js
assert(
  'P1P5P4CCCcP2P6P3'.match(reCriminals) &&
    'P1P5P4CCCcP2P6P3'.match(reCriminals)[0] == 'CCC'
);
```

你的正則表達式應該在字串 `P6P2P7P4P5CCCCCP3P1` 中匹配五個罪犯（`CCCCC`）。

```js
assert(
  'P6P2P7P4P5CCCCCP3P1'.match(reCriminals) &&
    'P6P2P7P4P5CCCCCP3P1'.match(reCriminals)[0] == 'CCCCC'
);
```

你的正規表達式不應該在空字串 `""` 中匹配任何罪犯。

```js
reCriminals.lastIndex = 0;
assert(!reCriminals.test(''));
```

你的正規表達式不應該匹配字串 `P1P2P3` 中的任何罪犯。

```js
reCriminals.lastIndex = 0;
assert(!reCriminals.test('P1P2P3'));
```

你的正規表達式應該在字串 `P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3` 中匹配五十個罪犯（`CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCC`）。

```js
assert(
  'P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3'.match(
    reCriminals
  ) &&
    'P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3'.match(
      reCriminals
    )[0] == 'CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCC'
);
```

# --seed--

## --seed-contents--

```js
let reCriminals = /./; // Change this line
```

# --solutions--

```js
let reCriminals = /C+/; // Change this line
```
