---
id: 587d7db9367417b2b2512ba6
title: 只指定較低的匹配數量
challengeType: 1
forumTopicId: 301366
dashedName: specify-only-the-lower-number-of-matches
---

# --description--

你可以使用大括號和數量限定符來指定樣式的下限和上限數量。有時你只想指定樣式的下限數量而不設上限。

若只指定較低的樣式數字，請保留第一個數字後接逗號。

例如，要只匹配字串 `hah` 中字母 `a` 至少出現 `3` 次，你的正規表達式會是 `/ha{3,}h/`。

```js
let A4 = "haaaah";
let A2 = "haah";
let A100 = "h" + "a".repeat(100) + "h";
let multipleA = /ha{3,}h/;
multipleA.test(A4);
multipleA.test(A2);
multipleA.test(A100);
```

依序，這三個 `test` 呼叫會傳回 `true`、`false` 和 `true`。

# --instructions--

將正則表達式 `haRegex` 更改為僅在有四個或更多字母 `z` 的情況下匹配字 `Hazzah`。

# --hints--

你的正則表達式應該使用大括號。

```js
assert(haRegex.source.match(/{.*?}/).length > 0);
```

你的正則表達式不應該匹配字串 `Hazzah`

```js
haRegex.lastIndex = 0;
assert(!haRegex.test('Hazzah'));
```

你的正則表達式不應該匹配字串 `Hazzzah`

```js
haRegex.lastIndex = 0;
assert(!haRegex.test('Hazzzah'));
```

你的正規表達式應該符合字串 `Hazzzzah`

```js
assert('Hazzzzah'.match(haRegex)[0].length === 8);
```

你的正規表達式應該符合字串 `Hazzzzzah`

```js
assert('Hazzzzzah'.match(haRegex)[0].length === 9);
```

你的正規表達式應該符合字串 `Hazzzzzzah`

```js
assert('Hazzzzzzah'.match(haRegex)[0].length === 10);
```

你的正規表達式應該匹配包含 30 個 `z` 的字串 `Hazzah`。

```js
assert('Hazzzzzzzzzzzzzzzzzzzzzzzzzzzzzzah'.match(haRegex)[0].length === 34);
```

# --seed--

## --seed-contents--

```js
let haStr = "Hazzzzah";
let haRegex = /change/; // Change this line
let result = haRegex.test(haStr);
```

# --solutions--

```js
let haStr = "Hazzzzah";
let haRegex = /Haz{4,}ah/; // Change this line
let result = haRegex.test(haStr);
```
