---
id: 587d7db8367417b2b2512ba0
title: 匹配除字母和數字以外的所有內容
challengeType: 1
forumTopicId: 301353
dashedName: match-everything-but-letters-and-numbers
---

# --description--

你已經知道你可以使用捷徑 `\w` 來匹配字母數字字元 `[A-Za-z0-9_]`。你可能想要搜尋的自然樣式是字母數字字元的相反。

你可以使用 `\W` 來搜尋 `\w` 的相反。請注意，相反的樣式使用大寫字母。這個捷徑與 `[^A-Za-z0-9_]` 相同。

```js
let shortHand = /\W/;
let numbers = "42%";
let sentence = "Coding!";
numbers.match(shortHand);
sentence.match(shortHand);
```

第一次 `match` 呼叫會傳回值 `["%"]`，第二次會傳回 `["!"]`。

# --instructions--

使用縮寫字元類別 `\W` 來計算各種引號和字串中非英數字元的數量。

# --hints--

你的正規表達式應該使用全域的旗標。

```js
assert(nonAlphabetRegex.global);
```

你的正規表達式應該在字串 `The five boxing wizards jump quickly.` 中找到 6 個非英數字元。

```js
assert(
  'The five boxing wizards jump quickly.'.match(nonAlphabetRegex).length == 6
);
```

你的正規表達式應該使用簡寫字元來匹配非字母數字的字元。

```js
assert(/\\W/.test(nonAlphabetRegex.source));
```

你的正規表達式應該在字串 `Pack my box with five dozen liquor jugs.` 中找到 8 個非英數字元。

```js
assert(
  'Pack my box with five dozen liquor jugs.'.match(nonAlphabetRegex).length == 8
);
```

你的正規表達式應該在字串 `How vexingly quick daft zebras jump!` 中找到 6 個非英數字元。

```js
assert(
  'How vexingly quick daft zebras jump!'.match(nonAlphabetRegex).length == 6
);
```

你的正規表達式應該在字串 `123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.` 中找到 12 個非英數字元。

```js
assert(
  '123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.'.match(nonAlphabetRegex)
    .length == 12
);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let nonAlphabetRegex = /change/; // Change this line
let result = quoteSample.match(nonAlphabetRegex).length;
```

# --solutions--

```js
let quoteSample = "The five boxing wizards_jump quickly.";
let nonAlphabetRegex = /\W/g; // Change this line
let result = quoteSample.match(nonAlphabetRegex).length;
```
