---
id: 587d7db7367417b2b2512b9e
title: 匹配結尾字串樣式
challengeType: 1
forumTopicId: 301352
dashedName: match-ending-string-patterns
---

# --description--

在上一個挑戰中，你學會了使用插入字元來搜尋字串開頭的樣式。也有一種方法可以搜尋字串結尾的樣式。

你可以使用正規表示法結尾的美元符號字元 `$` 來搜尋字串的結尾。

```js
let theEnding = "This is a never ending story";
let storyRegex = /story$/;
storyRegex.test(theEnding);
let noEnding = "Sometimes a story will have to end";
storyRegex.test(noEnding);
```

第一次 `test` 呼叫會傳回 `true`，而第二次會傳回 `false`。

# --instructions--

使用錨點字元（`$`）來匹配字串 `caboose` 在字串 `caboose` 的結尾。

# --hints--

你應該在你的正規表示式中使用美元符號 `$` 錨點來搜尋 `caboose`。

```js
assert(lastRegex.source == 'caboose$');
```

你的正規表達式不應該使用任何旗標。

```js
assert(lastRegex.flags == '');
```

你應該在字串 `The last car on a train is the caboose` 的結尾處比對 `caboose`。

```js
lastRegex.lastIndex = 0;
assert(lastRegex.test('The last car on a train is the caboose'));
```

# --seed--

## --seed-contents--

```js
let caboose = "The last car on a train is the caboose";
let lastRegex = /change/; // Change this line
let result = lastRegex.test(caboose);
```

# --solutions--

```js
let caboose = "The last car on a train is the caboose";
let lastRegex = /caboose$/; // Change this line
let result = lastRegex.test(caboose);
```
