---
id: 587d7dbb367417b2b2512bab
title: 使用捕獲群組來搜尋和取代
challengeType: 1
forumTopicId: 301368
dashedName: use-capture-groups-to-search-and-replace
---

# --description--

搜尋很有用。然而，當你同時變更（或取代）你匹配的文字時，你可以讓搜尋變得更強大。

你可以使用字串上的 `.replace()` 來搜尋並替換字串中的文字。`.replace()` 的輸入首先是你想搜尋的正規表達式樣式。第二個參數是用來替換符合項目的字串，或是一個用來執行某些操作的函式（程式）。

```js
let wrongText = "The sky is silver.";
let silverRegex = /silver/;
wrongText.replace(silverRegex, "blue");
```

`replace` 呼叫會傳回字串 `The sky is blue.`。

你也可以在取代字串中使用美元符號（`$`）來存取擷取群組。

```js
"Code Camp".replace(/(\w+)\s(\w+)/, '$2 $1');
```

`replace` 呼叫會傳回字串 `Camp Code`。

# --instructions--

撰寫一個使用三個捕獲群組的正規表達式 `fixRegex`，用來搜尋字串 `one two three` 中的每個字。然後更新 `replaceText` 變數，將 `one two three` 替換成字串 `three two one`，並將結果指定給 `result` 變數。請確保你在替換字串中使用捕獲群組，並使用美元符號（`$`）語法。

# --hints--

你應該使用 `.replace()` 來搜尋和取代。

```js
assert(__helpers.removeJSComments(code).match(/\.replace\(.*\)/));
```

你的正則表達式應該將字串 `one two three` 變更為字串 `three two one`

```js
assert(result === 'three two one');
```

你不應該更改最後一行。

```js
assert(__helpers.removeJSComments(code).match(/result\s*=\s*str\.replace\(.*?\)/));
```

`fixRegex` 應該至少使用三個捕獲群組。

```js
assert(new RegExp(fixRegex.source + '|').exec('').length - 1 >= 3);
```

`replaceText` 應該使用括號括起的子比對字串（例如，第 n 個括號括起的子比對字串 $n 對應第 n 個捕獲群組）。

```js
{
  const re = /(\$\d{1,2})+(?:[\D]|\b)/g;
  assert(replaceText.match(re).length >= 3);
}
```

# --seed--

## --seed-contents--

```js
let str = "one two three";
let fixRegex = /change/; // Change this line
let replaceText = ""; // Change this line
let result = str.replace(fixRegex, replaceText);
```

# --solutions--

```js
let str = "one two three";
let fixRegex = /(\w+) (\w+) (\w+)/g; // Change this line
let replaceText = "$3 $2 $1"; // Change this line
let result = str.replace(fixRegex, replaceText);
```
