---
id: 587d7dbb367417b2b2512bac
title: 移除開頭和結尾的空白
challengeType: 1
forumTopicId: 301362
dashedName: remove-whitespace-from-start-and-end
---

# --description--

有時字串周圍的空白字元是不需要的，但卻存在。對字串的典型處理程序是移除字串開頭和結尾的空白字元。

# --instructions--

撰寫正則表達式並使用適當的字串方法來移除字串開頭和結尾的空白字元。

**注意：** `String.prototype.trim()` 方法在這裡可行，但你需要使用正規表達式完成此挑戰。

# --hints--

`result` 應該要相等於字串 `Hello, World!`

```js
assert(result === 'Hello, World!');
```

你的解法不應該使用 `String.prototype.trim()` 方法。

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?trim/));
```

`result` 變數不應直接設定為字串

```js
assert(!__helpers.removeJSComments(code).match(/result\s*=\s*["'`].*?["'`]/));
```

`hello` 變數的值不應該被更改。

```js
assert(hello === '   Hello, World!  ');
```

# --seed--

## --seed-contents--

```js
let hello = "   Hello, World!  ";
let wsRegex = /change/; // Change this line
let result = hello; // Change this line
```

# --solutions--

```js
let hello = "   Hello, World!  ";
let wsRegex = /^(\s+)(.+[^\s])(\s+)$/;
let result = hello.replace(wsRegex, '$2');
```
