---
id: 587d7db5367417b2b2512b94
title: 使用萬用字元句點匹配任何內容
challengeType: 1
forumTopicId: 301348
dashedName: match-anything-with-wildcard-period
---

# --description--

有時候你不會（或不需要）知道你樣式中的確切字元。思考所有符合條件的字，例如拼寫錯誤，會花費很長時間。幸運的是，你可以使用萬用字元 `.` 來節省時間。

萬用字元 `.` 會匹配任何一個字元。萬用字元也稱為 `dot` 和 `period`。你可以像使用其他字元一樣使用萬用字元。在正規表示式中，例如，如果你想匹配 `hug`、`huh`、`hut` 和 `hum`，你可以使用正規表示式 `/hu./` 來匹配這四個字。

```js
let humStr = "I'll hum a song";
let hugStr = "Bear hug";
let huRegex = /hu./;
huRegex.test(humStr);
huRegex.test(hugStr);
```

這兩個 `test` 呼叫都會傳回 `true`。

# --instructions--

完成正規表示式 `unRegex`，使其能匹配字串 `run`、`sun`、`fun`、`pun`、`nun` 和 `bun`。你的正規表示式應該使用萬用字元。

# --hints--

你應該使用 `.test()` 方法。

```js
assert(__helpers.removeJSComments(code).match(/\.test\(.*\)/));
```

你應該在你的正規表達式 `unRegex` 中使用萬用字元字元。

```js
assert(/\./.test(unRegex.source));
```

你的正規表達式 `unRegex` 應該要在字串 `Let us go on a run.` 中符合 `run`。

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Let us go on a run.'));
```

你的正規表達式 `unRegex` 應該要在字串 `The sun is out today.` 中符合 `sun`。

```js
unRegex.lastIndex = 0;
assert(unRegex.test('The sun is out today.'));
```

你的正規表達式 `unRegex` 應該要在字串 `Coding is a lot of fun.` 中符合 `fun`。

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Coding is a lot of fun.'));
```

你的正則表達式 `unRegex` 應該要在字串 `Seven days without a pun makes one weak.` 中符合 `pun`。

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Seven days without a pun makes one weak.'));
```

你的正規表達式 `unRegex` 應該要在字串 `One takes a vow to be a nun.` 中符合 `nun`。

```js
unRegex.lastIndex = 0;
assert(unRegex.test('One takes a vow to be a nun.'));
```

你的正規表達式 `unRegex` 應該匹配字串 `She got fired from the hot dog stand for putting her hair in a bun.` 中的 `bun`。

```js
unRegex.lastIndex = 0;
assert(
  unRegex.test(
    'She got fired from the hot dog stand for putting her hair in a bun.'
  )
);
```

你的正規表達式 `unRegex` 不應該匹配字串 `There is a bug in my code.`。

```js
unRegex.lastIndex = 0;
assert(!unRegex.test('There is a bug in my code.'));
```

你的正規表達式 `unRegex` 不應該符合字串 `Catch me if you can.`。

```js
unRegex.lastIndex = 0;
assert(!unRegex.test('Catch me if you can.'));
```

# --seed--

## --seed-contents--

```js
let exampleStr = "Let's have fun with regular expressions!";
let unRegex = /change/; // Change this line
let result = unRegex.test(exampleStr);
```

# --solutions--

```js
let exampleStr = "Let's have fun with regular expressions!";
let unRegex = /.un/; // Change this line
let result = unRegex.test(exampleStr);
```
