---
id: 587d7db4367417b2b2512b93
title: 尋找多於第一個的匹配項
challengeType: 1
forumTopicId: 301342
dashedName: find-more-than-the-first-match
---

# --description--

到目前為止，你只能提取或搜尋一個樣式一次。

```js
let testStr = "Repeat, Repeat, Repeat";
let ourRegex = /Repeat/;
testStr.match(ourRegex);
```

這裡 `match` 會傳回 `["Repeat"]`。

若要搜尋或擷取樣式多於一次，你可以使用全域的搜尋旗標：`g`。

```js
let repeatRegex = /Repeat/g;
testStr.match(repeatRegex);
```

而這裡 `match` 傳回值為 `["Repeat", "Repeat", "Repeat"]`

# --instructions--

使用正規表達式 `starRegex`，從字串 `twinkleStar` 中尋找並擷取兩個 `Twinkle` 字。

**注意**你可以在你的正規表示式上有多個旗標，例如 `/search/gi`

# --hints--

你的正規表達式 `starRegex` 應該使用全域的旗標 `g`

```js
assert(starRegex.flags.match(/g/).length == 1);
```

你的正規表達式 `starRegex` 應該使用不區分大小寫的旗標 `i`

```js
assert(starRegex.flags.match(/i/).length == 1);
```

你的比對應該同時比對到字 `Twinkle` 的兩個出現位置。

```js
assert(
  result.sort().join() ==
    twinkleStar
      .match(/twinkle/gi)
      .sort()
      .join()
);
```

你的比對 `result` 應該包含兩個元素。

```js
assert(result.length == 2);
```

# --seed--

## --seed-contents--

```js
let twinkleStar = "Twinkle, twinkle, little star";
let starRegex = /change/; // Change this line
let result = twinkleStar; // Change this line
```

# --solutions--

```js
let twinkleStar = "Twinkle, twinkle, little star";
let starRegex = /twinkle/gi;
let result = twinkleStar.match(starRegex);
```
