---
id: 587d7db4367417b2b2512b92
title: 擷取符合項目
challengeType: 1
forumTopicId: 301340
dashedName: extract-matches
---

# --description--

到目前為止，你只是在檢查字串中是否存在某個樣式。你也可以使用 `.match()` 方法來擷取你找到的實際符合項目。

要使用 `.match()` 方法，請在字串上套用該方法，並將正規表達式放入小括弧內。

這裡有一個範例：

```js
"Hello, World!".match(/Hello/);
let ourStr = "Regular expressions";
let ourRegex = /expressions/;
ourStr.match(ourRegex);
```

這裡第一個 `match` 會傳回 `["Hello"]`，第二個會傳回 `["expressions"]`。

請注意，`.match` 語法是你迄今為止使用的 `.test` 方法的「相反」。

```js
'string'.match(/regex/);
/regex/.test('string');
```

# --instructions--

套用 `.match()` 方法來擷取字串 `coding`。

# --hints--

`result` 應該包含字串 `coding`

```js
assert(result.join() === 'coding');
```

你的正規表達式 `codingRegex` 應該搜尋字串 `coding`

```js
assert(codingRegex.source === 'coding');
```

你應該使用 `.match()` 方法。

```js
assert(__helpers.removeJSComments(code).match(/\.match\(.*\)/));
```

# --seed--

## --seed-contents--

```js
let extractStr = "Extract the word 'coding' from this string.";
let codingRegex = /change/; // Change this line
let result = extractStr; // Change this line
```

# --solutions--

```js
let extractStr = "Extract the word 'coding' from this string.";
let codingRegex = /coding/; // Change this line
let result = extractStr.match(codingRegex); // Change this line
```
