---
id: 587d7db9367417b2b2512ba5
title: 指定匹配的最大值和最小值
challengeType: 1
forumTopicId: 301367
dashedName: specify-upper-and-lower-number-of-matches
---

# --description--

請記住，你使用加號 `+` 來尋找一個或多個字元，使用星號 `*` 來尋找零個或多個字元。這些很方便，但有時你想要匹配某個特定範圍的樣式。

你可以使用 <dfn>數量指定符</dfn> 來指定樣式的下限和上限數量。數量指定符與大括號（`{` 和 `}`）一起使用。你將兩個數字放在大括號之間－分別表示樣式的下限和上限數量。

例如，要匹配只出現介於 `3` 和 `5` 次之間的字母 `a`，且出現在字串 `ah` 中，你的正規表達式會是 `/a{3,5}h/`。

```js
let A4 = "aaaah";
let A2 = "aah";
let multipleA = /a{3,5}h/;
multipleA.test(A4);
multipleA.test(A2);
```

第一次 `test` 呼叫會傳回 `true`，而第二次會傳回 `false`。

# --instructions--

將正則表達式 `ohRegex` 更改為僅在包含 3 到 6 個字母 h 時，匹配整個片語 `Oh no`。

# --hints--

你的正則表達式應該使用大括號。

```js
assert(ohRegex.source.match(/{.*?}/).length > 0);
```

你的正則表達式不應該匹配字串 `Ohh no`

```js
ohRegex.lastIndex = 0;
assert(!ohRegex.test('Ohh no'));
```

你的正規表達式應該符合字串 `Ohhh no`

```js
assert('Ohhh no'.match(ohRegex)[0].length === 7);
```

你的正則表達式應該符合字串 `Ohhhh no`

```js
assert('Ohhhh no'.match(ohRegex)[0].length === 8);
```

你的正規表達式應該符合字串 `Ohhhhh no`

```js
assert('Ohhhhh no'.match(ohRegex)[0].length === 9);
```

你的正規表達式應該符合字串 `Ohhhhhh no`

```js
assert('Ohhhhhh no'.match(ohRegex)[0].length === 10);
```

你的正則表達式不應該匹配字串 `Ohhhhhhh no`

```js
ohRegex.lastIndex = 0;
assert(!ohRegex.test('Ohhhhhhh no'));
```

# --seed--

## --seed-contents--

```js
let ohStr = "Ohhh no";
let ohRegex = /change/; // Change this line
let result = ohRegex.test(ohStr);
```

# --solutions--

```js
let ohStr = "Ohhh no";
let ohRegex = /Oh{3,6} no/; // Change this line
let result = ohRegex.test(ohStr);
```
