---
id: 587d7dba367417b2b2512ba8
title: 檢查全部或無任何項目
challengeType: 1
forumTopicId: 301338
dashedName: check-for-all-or-none
---

# --description--

有時你想要搜尋的樣式中可能包含存在與否不定的零件，但仍然可能需要檢查它們。

你可以用問號 `?` 指定元素可能存在。這會檢查前一個元素是否出現零次或一次。你可以將這個符號視為表示前一個元素是選用的。

例如，美式英語和英式英語有些微差異，你可以使用問號來匹配兩種拼法。

```js
let american = "color";
let british = "colour";
let rainbowRegex= /colou?r/;
rainbowRegex.test(american);
rainbowRegex.test(british);
```

`test` 方法的兩種用法都會傳回 `true`。

# --instructions--

將正則表達式 `favRegex` 更改為同時匹配美式英語（`favorite`）和英式英語（`favourite`）版本的字。

# --hints--

你的正規表示式應該使用選用符號 `?`。

```js
favRegex.lastIndex = 0;
assert(favRegex.source.match(/\?/).length > 0);
```

你的正規表達式應該符合字串 `favorite`

```js
favRegex.lastIndex = 0;
assert(favRegex.test('favorite'));
```

你的正規表達式應該符合字串 `favourite`

```js
favRegex.lastIndex = 0;
assert(favRegex.test('favourite'));
```

你的正則表達式不應該匹配字串 `fav`

```js
favRegex.lastIndex = 0;
assert(!favRegex.test('fav'));
```

# --seed--

## --seed-contents--

```js
let favWord = "favorite";
let favRegex = /change/; // Change this line
let result = favRegex.test(favWord);
```

# --solutions--

```js
let favWord = "favorite";
let favRegex = /favou?r/;
let result = favRegex.test(favWord);
```
