---
id: 587d7db8367417b2b2512ba3
title: 匹配空白字符
challengeType: 1
forumTopicId: 301359
dashedName: match-whitespace
---

# --description--

到目前為止，這些挑戰已涵蓋字母和數字的匹配。你也可以匹配字母之間的空白或空格。

你可以使用 `\s` 來搜尋空白字元，這是一個小寫的 `s`。這個樣式不僅符合空白字元，還包括回車、製表符、換頁符和換行字元。你可以將它視為類似字元類別 `[ \r\t\f\n\v]`。

```js
let whiteSpace = "Whitespace. Whitespace everywhere!"
let spaceRegex = /\s/g;
whiteSpace.match(spaceRegex);
```

這個 `match` 呼叫會傳回 `[" ", " "]`。
# --instructions--

將正則表達式 `countWhiteSpace` 更改為尋找字串中多個空白字元。

# --hints--

你的正規表達式應該使用全域的旗標。

```js
assert(countWhiteSpace.global);
```

你的正則表達式應該使用簡寫字元 `\s` 來匹配所有空白字元。

```js
assert(/\\s/.test(countWhiteSpace.source));
```

你的正規表達式應該在字串 `Men are from Mars and women are from Venus.` 中找到八個空格。

```js
assert(
  'Men are from Mars and women are from Venus.'.match(countWhiteSpace).length ==
    8
);
```

你的正規表達式應該在字串 `Space: the final frontier.` 中找到三個空格。

```js
assert('Space: the final frontier.'.match(countWhiteSpace).length == 3);
```

你的正則表達式應該在字串 `MindYourPersonalSpace` 中找不到空格。

```js
assert('MindYourPersonalSpace'.match(countWhiteSpace) == null);
```

# --seed--

## --seed-contents--

```js
let sample = "Whitespace is important in separating words";
let countWhiteSpace = /change/; // Change this line
let result = sample.match(countWhiteSpace);
```

# --solutions--

```js
let sample = "Whitespace is important in separating words";
let countWhiteSpace = /\s/g;
let result = sample.match(countWhiteSpace);
```
