---
id: 587d7db3367417b2b2512b8f
title: 匹配字面字串
challengeType: 1
forumTopicId: 301355
dashedName: match-literal-strings
---

# --description--

在上一個挑戰中，你使用正規表達式 `/Hello/` 搜尋字 `Hello`。該正規表達式搜尋字串 `Hello` 的字面匹配。這裡有另一個搜尋字串 `Kevin` 字面匹配的範例：

```js
let testStr = "Hello, my name is Kevin.";
let testRegex = /Kevin/;
testRegex.test(testStr);
```

這個 `test` 呼叫將傳回 `true`。

任何其他形式的 `Kevin` 都不會匹配。例如，正則表達式 `/Kevin/` 不會匹配 `kevin` 或 `KEVIN`。

```js
let wrongRegex = /kevin/;
wrongRegex.test(testStr);
```

這個 `test` 呼叫將傳回 `false`。

未來的挑戰將展示如何匹配那些其他表單。

# --instructions--

完成正規表達式 `waldoRegex`，以在字串 `waldoIsHiding` 中字面匹配尋找 `"Waldo"`。

# --hints--

你的正規表達式 `waldoRegex` 應該找到字串 `Waldo`

```js
waldoRegex.lastIndex = 0;
assert(waldoRegex.test(waldoIsHiding));
```

你的正則表達式 `waldoRegex` 不應該搜尋其他任何東西。

```js
waldoRegex.lastIndex = 0;
assert(!waldoRegex.test('Somewhere is hiding in this text.'));
```

你應該使用你的正則表達式執行字串的字面匹配。

```js
assert(!/\/.*\/i/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let waldoIsHiding = "Somewhere Waldo is hiding in this text.";
let waldoRegex = /search/; // Change this line
let result = waldoRegex.test(waldoIsHiding);
```

# --solutions--

```js
let waldoIsHiding = "Somewhere Waldo is hiding in this text.";
let waldoRegex = /Waldo/; // Change this line
let result = waldoRegex.test(waldoIsHiding);
```
