---
id: 587d7db9367417b2b2512ba7
title: 指定精確的匹配數量
challengeType: 1
forumTopicId: 301365
dashedName: specify-exact-number-of-matches
---

# --description--

你可以使用大括號和數量限定符來指定樣式的下限和上限。有時你只想要特定數量的匹配。

要指定特定數量的樣式，只需在大括號中放入該數字。

例如，要只匹配字 `hah` 中字母 `a` 出現 `3` 次，你的正規表達式會是 `/ha{3}h/`。

```js
let A4 = "haaaah";
let A3 = "haaah";
let A100 = "h" + "a".repeat(100) + "h";
let multipleHA = /ha{3}h/;
multipleHA.test(A4);
multipleHA.test(A3);
multipleHA.test(A100);
```

依序，三個 `test` 呼叫會傳回 `false`、`true` 和 `false`。

# --instructions--

將正則表達式 `timRegex` 更改為僅在有四個字母 `m` 時匹配字 `Timber`。

# --hints--

你的正則表達式應該使用大括號。

```js
assert(timRegex.source.match(/{.*?}/).length > 0);
```

你的正規表達式不應該匹配字串 `Timber`

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timber'));
```

你的正則表達式不應該匹配字串 `Timmber`

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timmber'));
```

你的正則表達式不應該匹配字串 `Timmmber`

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timmmber'));
```

你的正規表達式應該符合字串 `Timmmmber`

```js
timRegex.lastIndex = 0;
assert(timRegex.test('Timmmmber'));
```

你的正則表達式不應該匹配包含 30 個 `m` 的字串 `Timber`。

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Ti' + 'm'.repeat(30) + 'ber'));
```

# --seed--

## --seed-contents--

```js
let timStr = "Timmmmber";
let timRegex = /change/; // Change this line
let result = timRegex.test(timStr);
```

# --solutions--

```js
let timStr = "Timmmmber";
let timRegex = /Tim{4}ber/; // Change this line
let result = timRegex.test(timStr);
```
