---
id: 587d7db9367417b2b2512ba4
title: 匹配非空白字元
challengeType: 1
forumTopicId: 18210
dashedName: match-non-whitespace-characters
---

# --description--

你學會了使用小寫的 `s` 來搜尋空白字元的 `\s`。你也可以搜尋除了空白字元以外的所有內容。

使用 `\S` 來搜尋非空白字元，這是一個大寫的 `s`。此樣式不會符合空白字元、回車、製表符、換頁符和換行字元。你可以將它視為類似字元類別 `[^ \r\t\f\n\v]`。

```js
let whiteSpace = "Whitespace. Whitespace everywhere!"
let nonSpaceRegex = /\S/g;
whiteSpace.match(nonSpaceRegex).length;
```

`.length` 方法傳回的值會是 `32`。

# --instructions--

將正則表達式 `countNonWhiteSpace` 更改為尋找字串中多個非空白字元。

# --hints--

你的正規表達式應該使用全域的旗標。

```js
assert(countNonWhiteSpace.global);
```

你的正規表達式應該使用簡寫字元 `\S` 來匹配所有非空白字元。

```js
assert(/\\S/.test(countNonWhiteSpace.source));
```

你的正規表達式應該在字串 `Men are from Mars and women are from Venus.` 中找到 35 個非空白字元。

```js
assert(
  'Men are from Mars and women are from Venus.'.match(countNonWhiteSpace)
    .length == 35
);
```

你的正則表達式應該在字串 `Space: the final frontier.` 中找到 23 個非空白字元。

```js
assert('Space: the final frontier.'.match(countNonWhiteSpace).length == 23);
```

你的正則表達式應該在字串 `MindYourPersonalSpace` 中找到 21 個非空白字元。

```js
assert('MindYourPersonalSpace'.match(countNonWhiteSpace).length == 21);
```

# --seed--

## --seed-contents--

```js
let sample = "Whitespace is important in separating words";
let countNonWhiteSpace = /change/; // Change this line
let result = sample.match(countNonWhiteSpace);
```

# --solutions--

```js
let sample = "Whitespace is important in separating words";
let countNonWhiteSpace = /\S/g; // Change this line
let result = sample.match(countNonWhiteSpace);
```
