---
id: 587d7db8367417b2b2512ba2
title: 限制可能的使用者名稱
challengeType: 1
forumTopicId: 301363
dashedName: restrict-possible-usernames
---

# --description--

使用者名稱在整個網際網路上無處不在。它們賦予用戶在他們喜愛的網站上獨特的身份。

你需要檢查資料庫中所有的使用者名稱。以下是使用者在建立使用者名稱時必須遵守的一些簡單規則。

使用者名稱只能使用英數字元。

使用者名稱中唯一的數字必須位於結尾。結尾可以有零個或多個數字。使用者名稱不能以數字開頭。

使用者名稱字母可以是小寫和大寫。

使用者名稱必須至少有兩個字元長。兩個字元的使用者名稱只能使用英文字母作為字元。

# --instructions--

將正則表達式 `userCheck` 更改為符合上述條件的要求。

# --hints--

你的正規表達式應該符合字串 `JACK`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('JACK'));
```

你的正則表達式不應該匹配字串 `J`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('J'));
```

你的正規表達式應該符合字串 `Jo`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('Jo'));
```

你的正則表達式應該符合字串 `Oceans11`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('Oceans11'));
```

你的正規表達式應該符合字串 `RegexGuru`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('RegexGuru'));
```

你的正則表達式不應該匹配字串 `007`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('007'));
```

你的正則表達式不應該匹配字串 `9`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('9'));
```

你的正則表達式不應該匹配字串 `A1`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('A1'));
```

你的正則表達式不應該匹配字串 `BadUs3rnam3`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('BadUs3rnam3'));
```

你的正規表達式應該符合字串 `Z97`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('Z97'));
```

你的正規表達式不應該匹配字串 `c57bT3`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('c57bT3'));
```

你的正規表示式應該符合字串 `AB1`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('AB1'));
```

你的正則表達式不應該匹配字串 `J%4`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('J%4'))
```

# --seed--

## --seed-contents--

```js
let username = "JackOfAllTrades";
let userCheck = /change/; // Change this line
let result = userCheck.test(username);
```

# --solutions--

```js
let username = "JackOfAllTrades";
const userCheck = /^[a-z]([0-9]{2,}|[a-z]+\d*)$/i;
let result = userCheck.test(username);
```
