---
id: 587d7db8367417b2b2512ba1
title: 匹配所有非數字字符
challengeType: 1
forumTopicId: 301347
dashedName: match-all-non-numbers
---

# --description--

上一個挑戰示範了如何使用小寫的 `d` 的捷徑 `\d` 來搜尋數字。你也可以使用類似的捷徑，只是改用大寫的 `D`，來搜尋非數字。

尋找非數字字元的捷徑是 `\D`。這相當於字元類別 `[^0-9]`，該類別尋找一個不是介於零到九之間數字的單一字元。

# --instructions--

使用非數字的縮寫字元類別 `\D` 來計算電影標題中有多少非數字。

# --hints--

你的正則表達式應該使用捷徑字元來匹配非數字字元。

```js
assert(/\\D/.test(noNumRegex.source));
```

你的正規表達式應該使用全域的旗標。

```js
assert(noNumRegex.global);
```

你的正則表達式在字串 `9` 中不應找到非數字字元。

```js
assert('9'.match(noNumRegex) == null);
```

你的正規表達式應該在字串 `Catch 22` 中找到 6 個非數字。

```js
assert('Catch 22'.match(noNumRegex).length == 6);
```

你的正規表達式應該在字串 `101 Dalmatians` 中找到 11 個非數字。

```js
assert('101 Dalmatians'.match(noNumRegex).length == 11);
```

你的正則表達式應該在字串 `One, Two, Three` 中找到 15 個非數字。

```js
assert('One, Two, Three'.match(noNumRegex).length == 15);
```

你的正規表達式應該在字串 `21 Jump Street` 中找到 12 個非數字。

```js
assert('21 Jump Street'.match(noNumRegex).length == 12);
```

你的正規表達式應該在字串 `2001: A Space Odyssey` 中找到 17 個非數字。

```js
assert('2001: A Space Odyssey'.match(noNumRegex).length == 17);
```

# --seed--

## --seed-contents--

```js
let movieName = "2001: A Space Odyssey";
let noNumRegex = /change/; // Change this line
let result = movieName.match(noNumRegex).length;
```

# --solutions--

```js
let movieName = "2001: A Space Odyssey";
let noNumRegex = /\D/g; // Change this line
let result = movieName.match(noNumRegex).length;
```
