---
id: 587d7db3367417b2b2512b8e
title: 使用測試方法
challengeType: 1
forumTopicId: 301369
dashedName: using-the-test-method
---

# --description--

正規表達式用於程式語言中，以匹配字串的零件。你創建樣式來幫助你進行匹配。

如果你想在字串 `The dog chased the cat` 中找到字 `the`，你可以使用以下的正規表達式：`/the/`。請注意，正規表達式中不需要使用引號。

JavaScript 有多種使用正規表達式的方法。一種測試正規表達式的方法是使用 `.test()` 方法。`.test()` 方法會接受正規表達式，將它套用到字串（放在小括弧內），並且如果你的樣式有找到東西就傳回 `true`，否則傳回 `false`。

```js
let testStr = "freeCodeCamp";
let testRegex = /Code/;
testRegex.test(testStr);
```

這裡的 `test` 方法傳回 `true`。

# --instructions--

使用 `.test()` 方法將正則表達式 `myRegex` 套用到字串 `myString`。

# --hints--

你應該使用 `.test()` 來測試正則表達式。

```js
assert(__helpers.removeJSComments(code).match(/myRegex.test\(\s*myString\s*\)/));
```

你的結果應該傳回 `true`。

```js
assert(result === true);
```

# --seed--

## --seed-contents--

```js
let myString = "Hello, World!";
let myRegex = /Hello/;
let result = myRegex; // Change this line
```

# --solutions--

```js
let myString = "Hello, World!";
let myRegex = /Hello/;
let result = myRegex.test(myString); // Change this line
```
