---
id: 587d7db6367417b2b2512b9a
title: 匹配出現零次或多次的字元
challengeType: 1
forumTopicId: 301351
dashedName: match-characters-that-occur-zero-or-more-times
---

# --description--

上一個挑戰使用加號 `+` 來尋找出現一次或多次的字元。還有一個選項，可匹配出現零次或多次的字元。

執行此操作的字元是星號或星形：`*`。

```js
let soccerWord = "gooooooooal!";
let gPhrase = "gut feeling";
let oPhrase = "over the moon";
let goRegex = /go*/;
soccerWord.match(goRegex);
gPhrase.match(goRegex);
oPhrase.match(goRegex);
```

依序，三個 `match` 呼叫會傳回值 `["goooooooo"]`、`["g"]` 和 `null`。

# --instructions--

為你的 `chewieQuote` 生成一個正則表達式 `chewieRegex`，該正則表達式使用 `*` 字元來匹配緊接著零個或多個小寫 `a` 字元的大寫 `A` 字元。你的正則表達式不需要旗標或字元類別，且不應匹配任何其他引號。

# --hints--

你的正規表達式 `chewieRegex` 應該使用 `*` 字元來匹配零個或多個 `a` 字元。

```js
assert(/\*/.test(chewieRegex.source));
```

你的正規表達式應該匹配 `chewieQuote` 中的字串 `A`。

```js
assert(result[0][0] === 'A');
```

你的正規表達式應該匹配 `chewieQuote` 中的字串 `Aaaaaaaaaaaaaaaa`。

```js
assert(result[0] === 'Aaaaaaaaaaaaaaaa');
```

你的正規表示式 `chewieRegex` 應該要符合 `chewieQuote` 中的 16 個字元。

```js
assert(result[0].length === 16);
```

你的正則表達式不應該匹配字串 `He made a fair move. Screaming about it can't help you.` 中的任何字元。

```js
assert(
  !"He made a fair move. Screaming about it can't help you.".match(chewieRegex)
);
```

你的正規表達式不應該匹配字串 `Let him have it. It's not wise to upset a Wookiee.` 中的任何字元。

```js
assert(
  !"Let him have it. It's not wise to upset a Wookiee.".match(chewieRegex)
);
```

# --seed--

## --seed-contents--

```js
const chewieQuote = "Aaaaaaaaaaaaaaaarrrgh!";
// Only change code below this line
let chewieRegex = /change/; // Change this line
// Only change code above this line

let result = chewieQuote.match(chewieRegex);
```

# --solutions--

```js
const chewieQuote = "Aaaaaaaaaaaaaaaarrrgh!";
let chewieRegex = /Aa*/;
let result = chewieQuote.match(chewieRegex);
```
