---
id: 587d7db5367417b2b2512b96
title: 匹配字母表的字母
challengeType: 1
forumTopicId: 301354
dashedName: match-letters-of-the-alphabet
---

# --description--

你已經看到如何使用 <dfn>字元集</dfn> 來指定一個要匹配的字元群組，但當你需要匹配大量字元範圍（例如，字母表中的每個字母）時，這會需要大量輸入。幸運的是，有一個內建的特性可以讓這變得簡短且簡單。

在字元集合中，你可以使用連字號字元：`-`，來定義要比對的字元範圍。

例如，要匹配小寫字母 `a` 到 `e`，你會使用 `[a-e]`。

```js
let catStr = "cat";
let batStr = "bat";
let matStr = "mat";
let bgRegex = /[a-e]at/;
catStr.match(bgRegex);
batStr.match(bgRegex);
matStr.match(bgRegex);
```

依序，這三個 `match` 呼叫會傳回值 `["cat"]`、`["bat"]` 和 `null`。

# --instructions--

比對字串 `quoteSample` 中的所有字母。

**注意**：請確保同時匹配大寫和小寫字母。

# --hints--

你的正則表達式 `alphabetRegex` 應該匹配 35 項目。

```js
assert(result.length == 35);
```

你的正規表達式 `alphabetRegex` 應該使用全域的旗標。

```js
assert(alphabetRegex.flags.match(/g/).length == 1);
```

你的正規表達式 `alphabetRegex` 應該使用不區分大小寫的旗標。

```js
assert(alphabetRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The quick brown fox jumps over the lazy dog.";
let alphabetRegex = /change/; // Change this line
let result = alphabetRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "The quick brown fox jumps over the lazy dog.";
let alphabetRegex = /[a-z]/gi; // Change this line
let result = quoteSample.match(alphabetRegex); // Change this line
```
