---
id: 587d7db6367417b2b2512b98
title: 匹配未指定的單一字元
challengeType: 1
forumTopicId: 301358
dashedName: match-single-characters-not-specified
---

# --description--

到目前為止，你已經創建了一組你想要匹配的字元，但你也可以創建一組你不想匹配的字元。這種類型的字元集合稱為 <dfn>否定字元集合</dfn>。

要創建一個否定字元集，你需要在開啟括號後且不想匹配的字元之前放置一個脫字符號（`^`）。

例如，`/[^aeiou]/gi` 會匹配所有不是母音的字元。請注意，像 `.`、`!`、`[`、`@`、`/` 和空白字元都會被匹配——被否定的母音字元集合只排除母音字元。

# --instructions--

創建一個單一的正則表達式，匹配所有不是數字或母音的字元。記得在正則表達式中包含適當的旗標。

# --hints--

你的正則表達式 `myRegex` 應該匹配 9 項目。

```js
assert(result.length == 9);
```

你的正規表達式 `myRegex` 應該使用全域的旗標。

```js
assert(myRegex.flags.match(/g/).length == 1);
```

你的正規表達式 `myRegex` 應該使用不區分大小寫的旗標。

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "3 blind mice.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "3 blind mice.";
let myRegex = /[^0-9aeiou]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
