---
id: 587d7db7367417b2b2512b9f
title: 匹配所有字母和數字
challengeType: 1
forumTopicId: 301346
dashedName: match-all-letters-and-numbers
---

# --description--

使用字元類別，你能夠用 `[a-z]` 搜尋所有字母。這種類別字元足夠常見，因此有一個捷徑，雖然它也包含一些額外的字元。

JavaScript 中最接近匹配字母的字元類別是 `\w`。這個捷徑相當於 `[A-Za-z0-9_]`。這個字元類別匹配大寫和小寫字母以及數字。請注意，這個字元類別也包含底線字元（`_`）。

```js
let longHand = /[A-Za-z0-9_]+/;
let shortHand = /\w+/;
let numbers = "42";
let varNames = "important_var";
longHand.test(numbers);
shortHand.test(numbers);
longHand.test(varNames);
shortHand.test(varNames);
```

這四個 `test` 呼叫都會傳回 `true`。

這些快捷字元類別也稱為 <dfn>簡寫字元類別</dfn>。

# --instructions--

使用縮寫字元類別 `\w` 來計算各種引號和字串中的字母數字字元數量。

# --hints--

你的正規表達式應該使用全域的旗標。

```js
assert(alphabetRegexV2.global);
```

你的正規表達式應該使用簡寫字元 `\w` 來匹配所有字母數字字元。

```js
assert(/\\w/.test(alphabetRegexV2.source));
```

你的正則表達式應該在字串 `The five boxing wizards jump quickly.` 中找到 31 個字母數字字元。

```js
assert(
  'The five boxing wizards jump quickly.'.match(alphabetRegexV2).length === 31
);
```

你的正規表達式應該在字串 `Pack my box with five dozen liquor jugs.` 中找到 32 個字母數字字元。

```js
assert(
  'Pack my box with five dozen liquor jugs.'.match(alphabetRegexV2).length ===
    32
);
```

你的正規表達式應該在字串 `How vexingly quick daft zebras jump!` 中找到 30 個字母數字字元。

```js
assert(
  'How vexingly quick daft zebras jump!'.match(alphabetRegexV2).length === 30
);
```

你的正則表達式應該在字串 `123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.` 中找到 36 個字母數字字元。

```js
assert(
  '123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.'.match(alphabetRegexV2)
    .length === 36
);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let alphabetRegexV2 = /change/; // Change this line
let result = quoteSample.match(alphabetRegexV2).length;
```

# --solutions--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let alphabetRegexV2 = /\w/g; // Change this line
let result = quoteSample.match(alphabetRegexV2).length;
```
