---
id: 587d7dbb367417b2b2512baa
title: 使用捕獲群組重複利用樣式
challengeType: 1
forumTopicId: 301364
dashedName: reuse-patterns-using-capture-groups
---

# --description--

假設你想要匹配像下面這樣出現多次的字、單字。

```js
let repeatStr = "row row row your boat";
```

你可以使用 `/row row row/`，但如果你不知道重複的特定字、單字怎麼辦？<dfn>捕捉群組</dfn> 可用來尋找重複的子字串。

捕獲群組是透過將要捕獲的正規表達式樣式包在小括弧中來構成。在此情況下，目標是捕獲由字母數字字元組成的字，因此捕獲群組將是被小括弧包住的 `\w+`：`/(\w+)/`。

由群組所匹配的子字串會被存檔到一個暫時的「變數」，該變數可以在相同的正規表示式中使用反斜線和捕獲群組的編號（例如 `\1`）來存取。捕獲群組會依其開啟小括弧（由左至右）的位置自動編號，從 1 開始。

以下範例符合一個由空格分隔並出現三次的字。

```js
let repeatRegex = /(\w+) \1 \1/;
repeatRegex.test(repeatStr); // Returns true
repeatStr.match(repeatRegex); // Returns ["row row row", "row"]
```

在字串上使用 `.match()` 方法會傳回一個包含符合子字串及其捕獲群組的陣列。


# --instructions--

在 `reRegex` 中使用捕獲群組來匹配一個字串，該字串僅由相同的數字重複出現三次，且每個數字之間以單一空格分隔。

# --hints--

你的正規表示式應該使用數字的縮寫字元類別。

```js
assert(reRegex.source.match(/\\d/));
```

你的正則表達式應該重複使用一個捕獲群組兩次。

```js
assert(reRegex.source.match(/\\1|\\2/g).length >= 2);
```

你的正規表示式應該符合字串 `42 42 42`。

```js
reRegex.lastIndex = 0;
assert(reRegex.test('42 42 42'));
```

你的正規表達式應該符合字串 `100 100 100`。

```js
reRegex.lastIndex = 0;
assert(reRegex.test('100 100 100'));
```

你的正則表達式不應該匹配字串 `42 42 42 42`。

```js
assert.equal('42 42 42 42'.match(reRegex.source), null);
```

你的正則表達式不應該匹配字串 `42 42`。

```js
assert.equal('42 42'.match(reRegex.source), null);
```

你的正則表達式不應該匹配字串 `101 102 103`。

```js
reRegex.lastIndex = 0;
assert(!reRegex.test('101 102 103'));
```

你的正則表達式不應該匹配字串 `1 2 3`。

```js
reRegex.lastIndex = 0;
assert(!reRegex.test('1 2 3'));
```

你的正則表達式應該符合字串 `10 10 10`。

```js
reRegex.lastIndex = 0;
assert(reRegex.test('10 10 10'));
```

你的正則表達式不應該匹配字串 `42\t42\t42`。

```js
reRegex.lastIndex = 0;
assert(!reRegex.test('42\t42\t42'));
```

你的正則表達式不應該匹配字串 `42  42  42`。

```js
reRegex.lastIndex = 0;
assert(!reRegex.test('42  42  42'));
```

# --seed--

## --seed-contents--

```js
let repeatNum = "42 42 42";
let reRegex = /change/; // Change this line
let result = reRegex.test(repeatNum);
```

# --solutions--

```js
let repeatNum = "42 42 42";
let reRegex = /^(\d+) \1 \1$/;
let result = reRegex.test(repeatNum);
```
