---
id: 587d7db5367417b2b2512b97
title: 匹配數字和字母表字母
challengeType: 1
forumTopicId: 301356
dashedName: match-numbers-and-letters-of-the-alphabet
---

# --description--

使用連字號（`-`）來匹配字元範圍不僅限於字母，也適用於匹配數字範圍。

例如，`/[0-5]/` 會匹配介於 `0` 和 `5` 之間的任何數字，包括 `0` 和 `5`。

此外，也可以在單一字元集內結合字母和數字的範圍。

```js
let jennyStr = "Jenny8675309";
let myRegex = /[a-z0-9]/ig;
jennyStr.match(myRegex);
```

# --instructions--

創建一個單一的正則表達式，匹配介於 `h` 和 `s` 之間的字母範圍，以及介於 `2` 和 `6` 之間的數字範圍。請記得在正則表達式中包含適當的旗標。

# --hints--

你的正則表達式 `myRegex` 應該匹配 17 項目。

```js
assert(result.length == 17);
```

你的正規表達式 `myRegex` 應該使用全域的旗標。

```js
assert(myRegex.flags.match(/g/).length == 1);
```

你的正規表達式 `myRegex` 應該使用不區分大小寫的旗標。

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "Blueberry 3.141592653s are delicious.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "Blueberry 3.141592653s are delicious.";
let myRegex = /[h-s2-6]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
