---
id: 5c3dda8b4d8df89bea71600f
title: 檢查字元的混合群組
challengeType: 1
forumTopicId: 301339
dashedName: check-for-mixed-grouping-of-characters
---

# --description--

有時候我們想使用正規表達式來檢查字元群組，為此我們使用小括弧 `()`。

如果你想在字串中尋找 `Penguin` 或 `Pumpkin`，你可以使用以下的正規表達式：`/P(engu|umpk)in/g`

然後使用 `test()` 方法檢查測試字串中是否包含所需的字串群組。

```js
let testStr = "Pumpkin";
let testRegex = /P(engu|umpk)in/;
testRegex.test(testStr);
```

這裡的 `test` 方法會傳回 `true`。

# --instructions--

修正正則表達式，使其以區分大小寫的方式檢查 `Franklin Roosevelt` 或 `Eleanor Roosevelt` 的名稱，並且應該對中間名做出讓步。

然後修正程式碼，使你所創建的正規表達式能夠對 `myString` 進行檢查，並根據正規表達式是否匹配傳回 `true` 或 `false`。

# --hints--

你的正規表達式 `myRegex` 應該為字串 `Franklin D. Roosevelt` 傳回 `true`。

```js
myRegex.lastIndex = 0;
assert(myRegex.test('Franklin D. Roosevelt'));
```

你的正規表達式 `myRegex` 應該為字串 `Eleanor Roosevelt` 傳回 `true`。

```js
myRegex.lastIndex = 0;
assert(myRegex.test('Eleanor Roosevelt'));
```

你的正規表達式 `myRegex` 應該對字串 `Franklin Rosevelt` 傳回 `false`。

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('Franklin Rosevelt'));
```

你的正規表達式 `myRegex` 應該對字串 `Frank Roosevelt` 傳回 `false`。

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('Frank Roosevelt'));
```

你的正規表達式 `myRegex` 對字串 `FranklinRoosevelt` 應該傳回 `false`。

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('FranklinRoosevelt'));
```

你的正規表達式 `myRegex` 應該對字串 `EleanorRoosevelt` 傳回 `false`。

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('EleanorRoosevelt'));
```

你應該使用 `.test()` 來測試正則表達式。

```js
assert(__helpers.removeJSComments(code).match(/myRegex.test\(\s*myString\s*\)/));
```

你的結果應該傳回 `true`。

```js
assert(result === true);
```

# --seed--

## --seed-contents--

```js
let myString = "Eleanor Roosevelt";
let myRegex = /False/; // Change this line
let result = false; // Change this line
// After passing the challenge experiment with myString and see how the grouping works
```

# --solutions--

```js
let myString = "Eleanor Roosevelt";
let myRegex = /(Franklin|Eleanor) (([A-Z]\.?|[A-Z][a-z]+) )?Roosevelt/;
let result = myRegex.test(myString);
```
