---
id: 699b4e284291adcfcb90df47
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將練習使用原生型別以及物件和陣列型別，透過建置一個型別安全的使用者評測來達成。

首先建立一個名為 `profile` 的變數，並指定給它一個物件。你的物件應該包含以下三個屬性和對應的值：

```md
username: "codeLearner",
age: 25,
isLoggedIn: false,
```

# --hints--

你應該有一個名為 `profile` 的變數。

```js
assert.exists(profile);
```

你的 `profile` 變數應該是一個物件。

```js
assert.isObject(profile);
```

你的 `username` 屬性應該存在於 `profile` 物件中。

```js
assert.property(profile, "username");
```

你的 `username` 屬性應該有 `"codeLearner"` 的值。

```js
assert.propertyVal(profile, "username", "codeLearner");
```

你的 `profile` 物件應該有一個 `age` 屬性。

```js
assert.property(profile, "age");
```

你的 `age` 屬性應該有 `25` 的值。

```js
assert.propertyVal(profile, "age", 25);
```

你的 `profile` 物件應該有一個 `isLoggedIn` 屬性。

```js
assert.property(profile, "isLoggedIn");
```

你的 `isLoggedIn` 屬性應該有 `false` 的值。

```js
assert.propertyVal(profile, "isLoggedIn", false);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```
