---
id: 699b64096301fd75e3efeb1b
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

如果你再次打開主控台，你會看到不同的錯誤訊息：

```md
Property 'bio' is missing in type '{ username: string; age: number; isLoggedIn: false; }' but required in type '{ username: string; age: number; isLoggedIn: boolean; bio: string; }'.
```

如你在之前的課程中所學，你可以在屬性名稱旁使用 `?` 符號來標記該屬性為選用：

```ts
const musician: {
  name: string;
  instrument: string;
  isActive: boolean;
  numberOfAlbums?: number;
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

在這個範例中，`numberOfAlbums` 被標記為選用，因此 TypeScript 在這種情況下不會丟出錯誤。

將你的 `bio` 屬性更新為標記為選用。

# --hints--

你的 `bio` 屬性應該被標記為選用。

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string", isOptional: true };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
--fcc-editable-region--
  bio: string
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
