---
id: 699b6dd1a2961bd38491ef2e
title: 第 11 步驟
challengeType: 1
dashedName: step-11
---

# --description--

`profile` 物件有嚴格的鍵和值的型別。如果你嘗試新增未定義的屬性，或指定不相容的型別，TypeScript 就會丟出錯誤。

如果 `userRoles` 物件沒有那麼嚴格，並且擁有更彈性的鍵，那會很棒。

如你在之前的課程中所回顧的，你可以使用 `Record` 型別來達成這個結果：

```ts
const studentGrades: Record<string, number> = {
  Math: 95,
  English: 88,
  Science: 92
};
```

在這個 `studentGrades` 範例中，TypeScript 不需要事先知道所有的科目。這讓你可以稍後新增像 `Art` 或 `Music` 這樣的科目。TypeScript 只在意每個鍵的值必須是數字。

更新你的 `userRoles` 物件，為它添加一個行內 `Record` 型別。將它設為 `<string, string>`，使每個鍵可以是任意字串，而每個值必須是字串。

# --hints--

你的 `userRoles` 物件應該有一個行內 `Record` 型別，並設為 `Record<string, string>`。

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.annotation.matches("Record<string, string>"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles = {
--fcc-editable-region--
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

console.log(userRoles);
```
