---
id: 699b6d91cde15459d3619c4a
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

就像 `profile` 物件一樣，在主控台看到 `userRoles` 物件會很棒。

為了在主控台記錄 `console.log()` 物件，請添加 `userRoles`。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該將 `userRoles` 記錄到主控台。

```js
assert.equal(getLogs()[1], userRoles);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

--fcc-editable-region--

--fcc-editable-region--
```
