---
id: 699b61ca8c9cb9bc1968364e
title: 第 5 步驟
challengeType: 1
dashedName: step-5
---

# --description--

如果你打開主控台，應該會看到以下錯誤訊息：

```md
Object literal may only specify known properties, and 'mood' does not exist in type '{ username: string; age: number; isLoggedIn: boolean; }'.
```

在純 JavaScript 中，你可以隨意為 `profile` 物件添加任何屬性和值。但在 TypeScript 中，編譯器會檢查物件是否只包含你定義的屬性（這裡是 `username`、`age` 和 `isLoggedIn`）。由於 `mood` 沒有包含在型別中，TypeScript 會將它視為無效屬性並發出錯誤。

請從你的程式碼中移除 `mood: null`。

# --hints--

你的程式碼中不應該再有 `mood: null`。

```js
assert.isUndefined(profile.mood);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
--fcc-editable-region--
  mood: null
--fcc-editable-region--
};

console.log(profile);
```
