---
id: 699b65caf17a8e8d340af27d
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

現在是時候為 `profile` 物件添加最後一個屬性了。

在之前的課程中，你已經學會如何使用陣列型別。你也了解過不同的陣列型別，例如 `string[]` 或 `number[]`：

```ts
const scores: number[] = [95, 87, 100];
```

在這個範例中，`scores` 陣列只接受數字。如果你嘗試加入其他型別，例如 `string` 或 `null`，TypeScript 就會丟出錯誤。

更新你的 `profile` 物件，為它添加一個 `programmingLanguages` 屬性，型別為 `string[]`。在物件字面量中，將此屬性設為 `["JavaScript", "Python", "C++"]`。

# --hints--

你的 `programmingLanguages` 物件裡應該有一個 `profile` 屬性。

```js
assert.property(profile, "programmingLanguages");
```

你的 `programmingLanguages` 屬性應該有 `["JavaScript", "Python", "C++"]` 的值。

```js
assert.deepEqual(profile.programmingLanguages, ["JavaScript", "Python", "C++"]);
```

你的 `programmingLanguages` 屬性應該被定義為 `string[]` 型別。

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "programmingLanguages", type: "string[]" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};
--fcc-editable-region--

console.log(profile);
```
