---
id: 699b5b7cb68b6d35afca61b7
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

雖然可以繼續像這樣為 `profile` 物件添加屬性，但限制允許的屬性型別會更好。事先定義 `profile` 物件的結構會更理想。

在之前的課程中，你學過如何使用像這樣的物件型別：

```ts
const composer: { 
  name: string; 
  instrument: string; 
  isActive: boolean; 
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

在這個範例中，`composer` 是一個具有行內物件型別的物件。這個物件有三個必填屬性，分別是 `name`、`instrument` 和 `isActive`。前兩個屬性必須是 `string` 型別，`isActive` 必須是 `boolean` 型別。

更新你現有的 `profile` 物件，為它加入行內物件型別。`username` 應該被定義為 `string` 型別，`age` 應該被定義為 `number` 型別，`isLoggedIn` 應該被定義為 `boolean` 型別。

# --hints--

你的 `profile` 物件應該有一個行內物件型別，並且包含一個 `username` 屬性，其型別為 `string`。

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "username", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

你的 `profile` 物件應該有一個行內物件型別，並且包含一個 `age` 屬性，其型別為 `number`。

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "age", type: "number" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

你的 `profile` 物件應該有一個行內物件型別，並且包含一個 `isLoggedIn` 屬性，其型別為 `boolean`。

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "isLoggedIn", type: "boolean" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  mood: null
}
--fcc-editable-region--

console.log(profile);
```
