---
id: 699b71d98c820ac7ab9eea2a
title: 第 13 步驟
challengeType: 1
dashedName: step-13
---

# --description--

如果你再次打開主控台，你會看到一則新的錯誤訊息：

```md
Type 'number' is not assignable to type 'string'.
```

這是因為你嘗試為 `guest` 屬性賦予數字型別的值。

要移除 TypeScript 錯誤，請將 `guest` 屬性的值改為 `"read-only"`。

做完最後這個修改後，你的使用者評測工作坊就完成了！

# --hints--

你的 `guest` 屬性對於 `userRoles` 物件應該有 `"read-only"` 的值。

```js
assert.propertyVal(userRoles, "guest", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
--fcc-editable-region--
  guest: 3
--fcc-editable-region--
};

console.log(userRoles);
```

# --solutions--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
  guest: "read-only"
};

console.log(userRoles);
```
