---
id: 699b70b10ae344c802d106b0
title: 第 12 步
challengeType: 1
dashedName: step-12
---

# --description--

現在是時候為 `userRoles` 物件添加更多屬性了。

新增一個 `moderator` 屬性，值為 `"medium-access"`。

接著在下面新增一個 `guest` 屬性，值為 `3`。

# --hints--

你的 `userRoles` 物件應該有一個 `moderator` 屬性。

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.moderator);
```

你的 `moderator` 屬性應該有一個值為 `"medium-access"`。

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.moderator.value.matches("'medium-access'"));
```

你的 `userRoles` 物件應該有一個 `guest` 屬性。

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.guest);
```

你的 `guest` 屬性應該有一個數字型別的值 `3`。

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.guest.value.matches("3"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"

};
--fcc-editable-region--

console.log(userRoles);
```
