---
id: 699b62c99ec8c63470313365
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

現在是時候為你的 `profile` 物件添加另一個屬性了。

先從更新行內物件型別開始，加入一個 `bio` 屬性，其型別設定為 `string`。

# --hints--

你應該更新行內物件型別，加入一個 `bio` 屬性，並將其設定為 `string`。

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
--fcc-editable-region--
  username: string;
  age: number;
  isLoggedIn: boolean;
  
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
