---
id: 699b68f94a0d154c2a7a5ce0
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

現在 `profile` 物件已完成，是時候建置一個新的物件來表述使用者角色。

創建一個名為 `userRoles` 的新變數，並指定給它一個物件。你的物件應該包含以下屬性和值：

```md
admin: "full-access",
editor: "limited-access",
viewer: "read-only"
```

# --hints--

你應該有一個名為 `userRoles` 的變數。

```js
assert.isDefined(userRoles);
```

你的 `userRoles` 變數應該是一個物件。

```js
assert.isObject(userRoles);
```

你的 `userRoles` 物件應該有一個 `admin` 屬性。

```js
assert.property(userRoles, "admin");
```

你的 `admin` 屬性應該有 `"full-access"` 的值。

```js
assert.propertyVal(userRoles, "admin", "full-access");
```

你的 `userRoles` 物件應該有一個 `editor` 屬性。

```js
assert.property(userRoles, "editor");
```

你的 `editor` 屬性應該有 `"limited-access"` 的值。

```js
assert.propertyVal(userRoles, "editor", "limited-access");
```

你的 `userRoles` 物件應該有一個 `viewer` 屬性。

```js
assert.property(userRoles, "viewer");
```

你的 `viewer` 屬性應該有 `"read-only"` 的值。

```js
assert.propertyVal(userRoles, "viewer", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--

--fcc-editable-region--
```
