---
id: 699b5370a43e299b4673bcaa
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

在建置 `profile` 物件時，能在主控台看到它的屬性會很有幫助。

新增一個 `console.log`，將 `profile` 輸出到主控台。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該將 `profile` 輸出到主控台。

```js
assert.equal(getLogs()[0], profile);
```

# --seed--

## --seed-contents--

```ts
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
}

--fcc-editable-region--

--fcc-editable-region--
```
