---
id: 61537c5f81f0cf325b4a854c
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

أضف عنصر `header` داخل عنصر `body` وقم بتعيين صنف `header` له.

داخل `header`، أنشئ عنصر `h1` يحتوي على النص `css flexbox photo gallery`.

# --hints--

يجب أن يكون لديك عنصر `header` داخل عنصر `body`.

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

يجب أن يحتوي عنصر `header` على `class` بقيمة `header`.

```js
assert.isTrue(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

يجب أن يحتوي عنصر `header` على عنصر `h1` بداخله.

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

يجب أن يحتوي عنصر `h1` على النص `css flexbox photo gallery` بداخله.

```js
assert.equal(document?.querySelector('.header')?.querySelector('h1')?.innerText, 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
