---
id: 61537d86bdc3dd343688fceb
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

بعد ذلك، أعطِ كل عنصر `img` خاصية `src` حسب ترتيبه في المستند. يجب أن يحتوي أول عنصر `img` على خاصية `src` بقيمة `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`. يجب أن تكون بقية القيم مماثلة، مع استبدال الرقم `1` بالرقم الذي يمثل ترتيب عنصر `img` في المستند.

# --hints--

يجب أن تحتوي جميع عناصر `img` التسعة على خاصية `src`.

```js
const images = [...document.querySelectorAll('img')];
assert.isTrue(images.every(image => image.getAttribute('src')));
```

يجب أن يحتوي أول عنصر `img` على القيمة `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg` كقيمة لخاصية `src`.

```js
assert.equal(document.querySelectorAll('img')?.[0]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

يجب أن يحتوي ثاني عنصر `img` على القيمة `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg` كقيمة لخاصية `src`.

```js
assert.equal(document.querySelectorAll('img')?.[1]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg');
```

يجب أن يحتوي ثالث عنصر `img` على القيمة `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg` كقيمة لخاصية `src`.

```js
assert.equal(document.querySelectorAll('img')?.[2]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg');
```

يجب أن يحتوي رابع عنصر `img` على القيمة `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg` كقيمة لخاصية `src`.

```js
assert.equal(document.querySelectorAll('img')?.[3]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg');
```

يجب أن يحتوي خامس عنصر `img` على القيمة `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg` كقيمة لخاصية `src`.

```js
assert.equal(document.querySelectorAll('img')?.[4]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg');
```

يجب أن يحتوي سادس عنصر `img` على القيمة `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg` كقيمة لخاصية `src`.

```js
assert.equal(document.querySelectorAll('img')?.[5]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg');
```

يجب أن يحتوي سابع عنصر `img` على القيمة `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg` كقيمة لخاصية `src`.

```js
assert.equal(document.querySelectorAll('img')?.[6]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg');
```

يجب أن يحتوي ثامن عنصر `img` على القيمة `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg` كقيمة لخاصية `src`.

```js
assert.equal(document.querySelectorAll('img')?.[7]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg');
```

يجب أن يحتوي تاسع عنصر `img` على القيمة `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg` كقيمة لخاصية `src`.

```js
assert.equal(document.querySelectorAll('img')?.[8]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
--fcc-editable-region--
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
