---
id: 6153908a366afb4d57185c8d
title: الخطوة 13
challengeType: 0
dashedName: step-13
---

# --description--

لـ Flexbox محور رئيسي ومحور متقاطع. يُعرّف المحور الرئيسي بواسطة خاصية `flex-direction`، التي لها أربع قيم ممكنة:

- `row` (الافتراضي): محور أفقي مع عناصر مرنة من اليسار إلى اليمين
- `row-reverse`: محور أفقي مع عناصر مرنة من اليمين إلى اليسار
- `column`: محور عمودي مع عناصر مرنة من الأعلى إلى الأسفل
- `column-reverse`: محور عمودي مع عناصر مرنة من الأسفل إلى الأعلى

**ملاحظة**: المحاور والاتجاهات تختلف حسب اتجاه النص. القيم المعروضة هنا خاصة باتجاه النص من اليسار إلى اليمين.

جرّب القيم المختلفة لترى كيف تؤثر على التخطيط.

عندما تنتهي، عيّن خاصية `flex-direction` صريحة بقيمة `row` على عنصر `.gallery`.

# --hints--

يجب أن يحتوي محدِّد `.gallery` على خاصية `flex-direction` بقيمة `row`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.flexDirection, 'row');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}


.gallery {
  display: flex;
--fcc-editable-region--
  
--fcc-editable-region--
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
