---
id: 619665c9abd72906f3ad30f9
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Você estará construindo um feliz Flappy Penguin e praticando ainda mais transformações e animações CSS no processo.

Comece vinculando sua folha de estilos à página.

# --hints--

O código deve ter um elemento `link`.

```js
assert.exists(document.querySelector('link'));
```

O elemento `link` deve estar dentro do elemento `head`.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
