---
id: 6852dd8486c7ee0171322950
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

你需要一个变量来保存秘密代码以便解码，所以删除 `pass` 并在函数内部创建一个名为 `secret_code` 的变量，并赋值为空字串。

# --hints--

你应该在 `pin_extractor` 函数内有一个 `secret_code` 变量。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_code')
    `)
})
```

变量 `secret_code` 应该具有空字串的值。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_code').is_equivalent('secret_code = ""')
    `)
})
```

你的代码中不应再有 `pass`。

```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').has_pass()
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
