---
id: 6852db755c9c1274d75aae9a
title: 第 1 步
challengeType: 20
dashedName: step-1
---

# --description--

在本次工作坊中，你将创建一个 PIN 提取器。PIN 数字隐藏在诗歌的每一行中。

首先创建一个名为 `pin_extractor` 的函数，带有一个参数 `poem`。没有体的函数会引起误差，因此在函数内部添加 `pass` 作为临时占位符，使其有效。

# --hints--

你应该有一个名为 `pin_extractor` 的函数。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('pin_extractor')
    `)
})
```

`pin_extractor` 函数应有一个名为 `poem` 的参数。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poem')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
