---
id: 68530231c05b93324e071104
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

当你编写函数时，拥有一个用于调试目的的函数调用会很有用，所以在函数外部创建一个包含这首诗的 `poem` 变量：

```md
Stars and the moon
shine in the sky
white and bright
until the end of the night
```

你可以为此使用多行字串。

然后，使用 `poem` 变量作为参数调用函数。

# --hints--

你应该有一个名为 `poem` 的全局变量。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_variable('poem')
    `)
})
```

`poem` 变量的值应当相等于 `Stars and the moon\nshine in the sky\nwhite and bright\nuntil the end of the night`。

```js
({
    test: () => runPython(`
    pm = """Stars and the moon
    shine in the sky
    white and bright
    until the end of the night"""
    assert _Node(_code).find_variable('poem').is_equivalent(f'poem = """{pm}"""')
    `)
})
```

你应该使用 `poem` 作为参数调用函数 `pin_extractor`。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('pin_extractor(poem)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''

--fcc-editable-region--

--fcc-editable-region--
```
