---
id: 68531fbb934d735415053274
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

现在你已经添加了函数调用，你可以在终端中看到函数内部产生的任何误差，并且你始终可以使用 `print` 来查看变量的值。

每一行中都隐藏着密码的一个数字，所以在函数内部使用 `split` 方法将字串分割成一个行的列表。以换行字符（`\n`）分割行，并将得到的列表赋值给一个名为 `lines` 的变量。

# --hints--

你应该在你的 `pin_extractor` 函数中有一个名为 `lines` 的变量。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('lines')
    `)
})
```

变量 `lines` 应该具有值 `poem.split('\n')`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('lines').is_equivalent('lines = poem.split("\\\\n")')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
```
