---
id: 685535bf527b2c77236b3a9a
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

第 n 个数字被隐藏为第 n 行第 n 个单词的长度。

要查找第 n 个单词的长度，下一步是将诗句分割成单词列表：在循环内，创建一个变量 `words` 并赋值为使用 `split` 方法将 `line` 拆分成单词的结果。

然后，为你的代码添加一个以 `words` 作为参数的 `print` 调用。

# --hints--

你应该在 `loop` 内有一个 `words` 变量。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_variable('words')
    `)
})
```

`words` 变量的值应为 `line.split()`。

```js
({
    test: () => runPython(`
    variable = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_variable('words')
    assert variable.is_equivalent('words = line.split()') or variable.is_equivalent('words = line.split(" ")')
    `)
})
```

你应该有 `print(words)`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        print(line)
--fcc-editable-region--
        
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
