---
id: 68553423d775cc60134ac8bb
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

你需要独立处理每一行：创建一个针对 `lines` 的循环，使用 `line` 作为循环变量。在循环内部，打印 `line`。

# --hints--

你的 `pin_extractor` 函数中应该有一个 `for` 循环。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0]
    `)
})
```

循环应迭代 `lines`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('lines')
    `)
})
```

循环应以 `line` 作为循环变量（`for line in lines`）。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line')
    `)
})
```

你应该在 `loop` 内部有 `print(line)`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
