---
id: 685578aae65e4106ad2b4300
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

正如你在之前的课程中学到的，`enumerate` 函数允许在遍历可迭代对象时跟踪索引：

```py
fruits = ['apple', 'plum', 'bananas']

for index, item in enumerate(fruits):
    print(index, item)

# 0 apple
# 1 plum
# 2 bananas
```

秘密代码的第 n 位数字来自第 n 行的第 n 个单词，所以你需要知道每一行的编号。

更改循环，使其在 `enumerate(lines)` 上迭代，并在 `line` 之前添加另一个名为 `line_index` 的循环变量。

另外，更新 `print` 调用为 `print(line_index, line)`。

# --hints--

循环应迭代 `enumerate(lines)`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('enumerate(lines)')
    `)
})
```

循环应以 `line_index, line` 作为循环变量。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line_index, line')
    `)
})
```

你应该在 `loop` 内部有 `print(line_index, line)`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line_index, line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    for line in lines:
        print(line)
--fcc-editable-region--
        words = line.split()
        print(words)

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
