---
id: 6855918e87c33c6979d21d3f
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

现在这个函数在从诗歌中提取代码方面表现良好，但有一个边缘分支需要考虑：诗歌的行可能比预期的短，这会导致函数出错（你可以尝试从诗歌中删除单词 `bright` 来验证）。

将语句 `secret_code += str(len(words[line_index]))` 放入一个 `if` 语句中，该语句检查 `words` 列表中是否有足够的单词。

# --hints--

在 `loop` 内你应该有一个 `if` 语句。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0]
    `)
})
```

if 语句条件应为 `len(words) > line_index`。


```js
({
    test: () => runPython(`
    _cond = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
    assert _cond.is_equivalent('len(words) > line_index') or _cond.is_equivalent('line_index < len(words)')
    `)
})
```

在 `if` 语句内你应该有 `secret_code += str(len(words[line_index]))`。


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt('secret_code += str(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
--fcc-editable-region--
        secret_code += str(len(words[line_index]))
--fcc-editable-region--
    return secret_code

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

print(pin_extractor(poem))

```
