---
id: 68592730aaf34b61ae279168
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

从 `len` 返回的长度是一个整数，你需要使用 `str` 函数将其转换为字串，才能将它添加到 `secret_code` 字串中。

更新 `print` 调用以将长度打印为字串。

# --hints--

你应该在循环末尾有 `print(str(len(words[line_index])))`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(str(len(words[line_index])))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
--fcc-editable-region--
        print(len(words[line_index]))
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
