---
id: 6859a082915dae1e0eaaf962
title: 步骤 19
challengeType: 20
dashedName: step-19
---

# --description--

最后，返回 `secret_codes`，然后取消注释函数调用并将参数更新为 `[poem, poem2, poem3]`。现在你的函数可以同时处理多首诗歌，并且能够提取多个密码！

`pin_extractor` 函数已完成。

# --hints--

函数应在结束时返回 `secret_codes`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_return('secret_codes')
    `)
})
```

你应该调用函数 `print(pin_extractor([poem, poem2, poem3]))`。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('print(pin_extractor([poem, poem2, poem3]))')
    `)
})
```

`pin_extractor([poem, poem2, poem3])` 应该返回 `['5202', '3346', '50000']`。

```js
({
    test: () => runPython(`
    assert pin_extractor([poem, poem2, poem3]) == ['5202', '3346', '50000']
    `)
})
```

该函数应能正确处理其他诗歌列表。

```js
({
    test: () => runPython(`
plum = '''In dusky shade the plum tree grew,
Its branches bent with evening dew.
A single fruit, so ripe and round,
Falls softly to the waiting ground.'''
apple = '''Red or green on branches high,
Like lanterns hung against the sky.
Crisp with autumn's golden breath,
A bite that wakes the world from death.'''
grape = '''Clustered tight on winding vines,
Bathed in sun where silence shines.
Jewel-toned beads in twilight air,
A whispered promise hanging there.'''
assert pin_extractor([plum, apple, grape]) == ['2863', '3885', '9227']
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poems):
    secret_codes = []
    for poem in poems:
        secret_code = ''
        lines = poem.split('\n')
        for line_index, line in enumerate(lines):
            words = line.split()
            if len(words) > line_index:
                secret_code += str(len(words[line_index]))
            else:
                secret_code += '0'
        secret_codes.append(secret_code)
--fcc-editable-region--
        

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'
poem3 = 'There\nonce\nwas\na\ndragon'

# print(pin_extractor(poem))
--fcc-editable-region--

```

# --solutions--

```py
def pin_extractor(poems):
    secret_codes = []
    for poem in poems:
        secret_code = ''
        lines = poem.split('\n')
        for line_index, line in enumerate(lines):
            words = line.split()
            if len(words) > line_index:
                secret_code += str(len(words[line_index]))
            else:
                secret_code += '0'
        secret_codes.append(secret_code)
    return secret_codes

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'
poem3 = 'There\nonce\nwas\na\ndragon'

print(pin_extractor([poem, poem2, poem3]))

```

