---
id: 685929ee6d500c6f2bce6751
title: 步骤 13
challengeType: 20
dashedName: step-13
---

# --description--

将 `pin_extractor(poem)` 传递给 `print` 函数，这样你就可以看到该函数返回的内容。

# --hints--

你应该调用 `print(pin_extractor(poem))`。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('print(pin_extractor(poem))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        secret_code += str(len(words[line_index]))
    return secret_code


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

--fcc-editable-region--
pin_extractor(poem)
--fcc-editable-region--

```
