---
id: 6859a037b3d6721c7fbb8055
title: 步骤18
challengeType: 20
dashedName: step-18
---

# --description--

在循环之前，创建一个初始值为空列表的 `secret_codes` 新变量。

然后将 `return secret_code` 行替换为一行，将 `secret_code` 追加到 `secret_codes` 列表。

# --hints--

你应该在 `pin_extractor` 函数内有一个 `secret_codes` 变量。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_codes')
    `)
})
```

变量 `secret_codes` 应该具有一个空列表的值。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_codes').is_equivalent('secret_codes = []')
    `)
})
```

在循环结束时你应该有 `secret_codes.append(secret_code)`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('secret_codes.append(secret_code)')
    `)
})
```

你不应该返回 `secret_code`。


```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_return('secret_code')
    `)
})
```


# --seed--

## --seed-contents--

```py
def pin_extractor(poems):
--fcc-editable-region--

    for poem in poems:
        secret_code = ''
        lines = poem.split('\n')
        for line_index, line in enumerate(lines):
            words = line.split()
            if len(words) > line_index:
                secret_code += str(len(words[line_index]))
            else:
                secret_code += '0'
        return secret_code
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'
poem3 = 'There\nonce\nwas\na\ndragon'

# print(pin_extractor(poem))

```
