---
id: 6859290fc8e07e6b2a8d8523
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

删除两个 `print` 调用，并在函数的最后一行添加一个返回 `secret_code` 的 `return` 语句。

# --hints--

你不应该有 `print` 调用。

```js
({
    test: () => runPython(`
    assert not _Node(_code).block_has_call('print')
    `)
})
```

`pin_extractor` 函数应在最后返回 `secret_code`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_return('secret_code')
    `)
})
```

`pin_extractor(poem)` 应该返回 `5262`。


```js
({
    test: () => runPython(`
    assert pin_extractor(poem) == '5262'
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
--fcc-editable-region--
        print(line_index, line)
        words = line.split()
        print(str(len(words[line_index])))
        secret_code += str(len(words[line_index]))
    
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
