---
id: 694648acde178bb8202d9513
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

學生姓名應該遵循與其他細節相同的格式。

移除 `name` 變數先前的輸出。然後，將 `name` 和 `type(name)` 用逗號分隔並排在同一行列印，就像前一步驟一樣。

# --hints--

你不應該在你的程式碼中有 `print(name)`。

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(name)")`)})
```

你的程式碼中不應該有 `print(type(name))`。

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(type(name))")`)})
```

你應該在同一次 `print()` 呼叫中使用逗號分隔符列印 `name` 和 `type(name)`。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name, type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--
print(name)
print(type(name))
--fcc-editable-region--

is_student = True
print(is_student, type(is_student))
```
