---
id: 694648acde178bb8202d9517
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

輸出是 `True`，確認 `score` 是 `float`。

使用單一的 `print()` 述語列印 `score` 值及其資料型別，以完成成績單。

# --hints--

你應該使用逗號分隔符列印 `score` 和 `type(score)`。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(score, type(score))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
--fcc-editable-region--

--fcc-editable-region--

```

# --solutions--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
print(score, type(score))
```
