---
id: 694648acde178bb8202d9515
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

現在，為學生添加分數。

宣告一個名為 `score` 的變數並指定其值為 `80.5`。

雖然 `age` 和 `score` 都是數字，但它們可能不是同一種類。Python 提供了一個名為 `isinstance()` 的函式來檢查這一點。

```py
x = 10
print(isinstance(x, int)) # Output: True
```

使用 `isinstance()` 來檢查 `score` 是否為 `int`，並將結果如上例所示列印到終端機。

# --hints--

你應該有一個 `score` 變數。

```js
({ test: () => runPython(`assert _Node(_code).has_variable("score")`) })
```

變數 `score` 應該儲存數值 `80.5`。不要用引號包住該數值。

```js
({ test: () => runPython(`assert _Node(_code).find_variable("score").is_equivalent("score = 80.5")`) })
```

你應該列印 `isinstance(score, int)` 的結果。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(isinstance(score, int))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

--fcc-editable-region--

--fcc-editable-region--

```
