---
id: 694648acde178bb8202d9512
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

成績單也應該顯示學生目前是否註冊。這可以使用布林值來表述。

布林值表述一個是或否的條件，並且它們經常用於在程式碼中做出決策。布林值只有兩個：`True` 和 `False`。

宣告一個名為 `is_student` 的變數並指定其值為 `True`。

# --hints--

你應該有一個 `is_student` 變數。

```js
({ test: () => runPython(`assert _Node(_code).has_variable("is_student")`) })
```

變數 `is_student` 應該儲存值 `True`。不要用引號包住該值。

```js
({ test: () => runPython(`assert _Node(_code).find_variable("is_student").is_equivalent("is_student = True")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

--fcc-editable-region--

--fcc-editable-region--
```
