---
id: 694648abde178bb8202d9511
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

你現在應該會在終端機中看到學生姓名被列印出來。

Python 提供了一個名為 `type()` 的函式，你可以用來檢查一個值的型別。

```py
platform = 'freeCodeCamp'
print(type(platform)) # Output: <class 'str'>
```

在上述範例中，輸出 `<class 'str'>` 表示傳遞給 `type()` 函式的變數是字串。

使用帶有 `name` 作為引數的 `type()` 函式，並像範例一樣列印輸出。在終端機中檢查輸出，顯示 `name` 的型別是 `str`（字串）。

# --hints--

你應該列印呼叫 `type(name)` 的結果。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
--fcc-editable-region--

--fcc-editable-region--
```
