---
id: 694648acde178bb8202d9516
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

輸出是 `False`，這表示 `score` 不是 `int`。

Python 中另一種常見的數字類型是 `float`，它表述帶有小數的數字。將現有 `isinstance()` 呼叫中的 `int` 替換為 `float` 以確認這一點。

# --hints--

你應該在現有的 `isinstance(score, int)` 呼叫中將 `int` 取代為 `float`。

```js
({ test: () => runPython(`
  assert not _Node(_code).has_call("print(isinstance(score, int))")
  assert _Node(_code).has_call("print(isinstance(score, float))")
`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
--fcc-editable-region--
print(isinstance(score, int))
--fcc-editable-region--

```
