---
id: 69461e4f98594012939bc1c6
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個工作坊中，你將透過建置一個簡單的成績單印表機來練習 Python 資料型別。

如你在之前的課程中所學，變數是以這種方式指定、設值：

```py
greeting = 'Hello World'
```

在上述範例中，變數 `greeting` 的值是一個字串，該字串是一連串被引號包圍的字元。

創建一個名為 `name` 的變數並指定其值為字串 `Alice`。請記得像範例中所示，用單引號或雙引號將值包起來。

**重要：** Python 使用縮排（行首的空格）來組織程式碼。對於本次工作坊，請確保每行程式碼開頭沒有多餘的空格。添加多餘的空格會導致 `IndentationError` 並阻止你的程式碼執行。

# --hints--

你應該有一個 `name` 變數。

```js
({ test: () => runPython(`assert _Node(_code).has_variable('name')`) })
```

變數 `name` 應該儲存值 `'Alice'`。

```js
({ test: () => runPython(`assert _Node(_code).find_variable("name").is_equivalent("name = 'Alice'")`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
