---
id: 694648acde178bb8202d9514
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

現在將學生的年齡添加到成績單資料中。為此，你將使用整數，這是 Python 中的數值型別之一。

宣告一個名為 `age` 的變數並指定其整數值為 `20`。

接著，將 `age` 的值和型別用逗號分隔一起列印。查看終端機的輸出，會顯示 `age` 的值和它的型別 `int`（整數）。

# --hints--

你應該有一個 `age` 變數。

```js
({ test: () => runPython(`assert _Node(_code).has_variable("age")`) })
```

變數 `age` 應該儲存值 `20`。不要用引號包住該值。

```js
({ test: () => runPython(`assert _Node(_code).find_variable("age").is_equivalent("age = 20")`) })
```

你應該使用逗號分隔符列印 `age` 和 `type(age)`。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(age, type(age))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

--fcc-editable-region--

--fcc-editable-region--

```
