---
id: 6947a0c2b6de035dd8d3eea1
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

你可以使用 `print()` 函式列印變數的值。

```py
greeting = 'Hello World'
print(greeting) # Output: Hello World
```

你將在接下來的課程中學習更多關於函式的內容。現在，請知道函式是一個可重複使用的程式碼區段，可以被呼叫或喚起來執行其程式碼，並且可以傳遞引數給它。

在上述範例中，`print(greeting)` 是一個函式呼叫，而 `greeting` 是該函式的引數。

參考範例並列印 `name` 變數。在終端機中檢查輸出。

# --hints--

你應該以 `name` 作為引數來呼叫 `print()`。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name)")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--

--fcc-editable-region--
```
