---
id: 6947a3e516bb8a4a681274a7
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

`print()` 函式（程式）可以同時顯示多個值。用逗號（`,`）分隔值，以便在同一行列印它們。

```py
subject = 'Python'
print(subject, type(subject)) # Output: Python <class 'str'>
```

在同一行使用逗號 `,` 列印 `is_student` 和 `type(is_student)`，如範例所示。然後，在終端機中檢查輸出，該輸出會顯示 `is_student` 的值及其型別為 `bool`（布林值）。

# --hints--

你應該使用逗號分隔列印 `is_student` 和 `type(is_student)`。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(is_student, type(is_student))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

is_student = True
--fcc-editable-region--

--fcc-editable-region--
```
