---
id: 6909ed80e1c195f49564edac
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將學習如何使用條件述語和比較運算子。

首先宣告一個名為 `truthyOrFalsy` 的變數，並將布林值 `true` 指定給它。

# --hints--

你應該有一個名為 `truthyOrFalsy` 的變數。

```js
assert.exists(truthyOrFalsy);
```

你應該將 `true` 指定給你的 `truthyOrFalsy` 變數。

```js
assert.isTrue(truthyOrFalsy);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
