---
id: 690a58cc4f828a49eab1b53a
title: 步驟 6
challengeType: 1
dashedName: step-6
---

# --description--

宣告一個名為 `hasDeveloperJob` 的變數，並將布林值 `true` 指定給它。

# --hints--

你應該有一個名為 `hasDeveloperJob` 的變數。

```js
assert.exists(hasDeveloperJob);
```

你應該將 `true` 指定到你的 `hasDeveloperJob` 變數。

```js
assert.isTrue(hasDeveloperJob);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
