---
id: 690a3ad44874a928a2fe3077
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

比較運算子允許你比較數值並撰寫更複雜的條件述語。

```js
if (today == birthday) {
  console.log("Happy Birthday!");
}
```

例如，只有當變數 `today` 和 `birthday` 的值相等時，上述程式碼才會將 `"Happy Birthday!"` 記錄到主控台。

創建第三個 `if` 述語。對於它的條件，使用 `>=` 運算子來檢查 `timmyAge` 是否大於或相等於 `16`。

在你新的 `if` 述語主體內，將 `"Timmy is old enough to drive."` 記錄到主控台。

# --hints--

你應該在你的程式碼中有第三個 `if` 述語。

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(.+\)\s*\{.+\}/s);
```

你的第三個 `if` 述語應該使用 `timmyAge >= 16` 作為其條件。

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{.+\}/s);
```

你應該在你的第三個 `if` 述語中將 `"Timmy is old enough to drive."` 記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 18;

--fcc-editable-region--

--fcc-editable-region--
```
