---
id: 690a57095c4afb27bbe751b3
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

Falsy 值是在布林值背景關係中計算結果為 `false` 的值。Falsy 值的範例有：

- `""` (空字串)
- `0`
- `false`
- `null`
- `undefined`
- `NaN`

你將會在這個課程中遇到所有這些。

現在，透過將 `truthyOrFalsy` 的值指定為空字串來變更指派。之後，你會看到 `false` 被記錄到主控台。

# --hints--

你應該將空字串指定給 `truthyOrFalsy`。

```js
assert.strictEqual(truthyOrFalsy, "");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "freeCodeCamp";
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
