---
id: 690a0f7a5c22c15222a6aa5c
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

現在創建另一個使用 `isTimmyAGamer` 作為條件的 `if` 述語。在你的 `if` 述語主體內，將 `"Timmy loves to play World of Warcraft."` 記錄到主控台。

請注意，因為你新的 `if` 述語的條件為假，所以你不會看到任何記錄到主控台的內容。

# --hints--

你的程式碼中應該有兩個 `if` 述語。

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(.+\)\s*\{.+\}/s);
```

你的第二個 `if` 述語應該使用 `isTimmyAGamer` 作為條件。

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{.+\}/s);
```

你應該在你的第二個 `if` 述語內將 `"Timmy loves to play World of Warcraft."` 記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy loves to play World of Warcraft.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

--fcc-editable-region--

--fcc-editable-region--
```
