---
id: 690a562fd611e8155599ce98
title: 步驟 3
challengeType: 1
dashedName: step-3
---

# --description--

透過將 `truthyOrFalsy` 的值指定為 `"freeCodeCamp"` 來變更指派。你仍然會看到 `true` 被記錄到主控台，因為非空字串是 truthy 值。

# --hints--

你的 `truthyOrFalsy` 變數應該有值 `"freeCodeCamp"`。

```js
assert.strictEqual(truthyOrFalsy, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = true;
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
