---
id: 690a06833d3e83b5226ade78
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

條件述語用於根據條件執行程式碼。它由 `if` 關鍵字組成，後面接著用小括弧括起的條件，以及用大括弧括起的主體。以下是一個 `if` 述語的範例：

```js
if (condition) {
  console.log("condition is truthy");
}
```

當 `condition` 是真值（非空字串、非零數字、`true` 等）時，會執行 `if` 述語主體內的程式碼。

在你的變數下方，創建一個使用 `hasDeveloperJob` 作為條件的 `if` 述語。在你的 `if` 述語主體內，將 `"Timmy is employed as a developer."` 記錄到主控台。

# --hints--

你應該有一個 `if` 述語。

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(.+\)\s*\{?/);
```

你的 `if` 述語應該使用 `hasDeveloperJob` 作為其條件。

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{?/);
```

你應該在你的 `if` 述語中將 `"Timmy is employed as a developer."` 記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is employed as a developer.\1\s*\)\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;
--fcc-editable-region--

--fcc-editable-region--
```
