---
id: 690a110a76d0496d9e7c28eb
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

宣告另一個名為 `timmyAge` 的變數並將其設為 `18`。

# --hints--

你應該有一個名為 `timmyAge` 的變數。

```js
assert.exists(timmyAge);
```

你應該將 `18` 指定到你的 `timmyAge` 變數。

```js
assert.strictEqual(timmyAge, 18);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

--fcc-editable-region--

--fcc-editable-region--
```
