---
id: 690a0df27ee9ee36d935d530
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

宣告另一個名為 `isTimmyAGamer` 的變數，並將其指定為布林值 `false`。

# --hints--

你應該宣告一個名為 `isTimmyAGamer` 的變數。

```js
assert.exists(isTimmyAGamer);
```

你應該將布林值 `false` 指定給你的 `isTimmyAGamer` 變數。

```js
assert.isFalse(isTimmyAGamer);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}
--fcc-editable-region--

--fcc-editable-region--
```
