---
id: 690a51806fb743bc896c3af7
title: 步驟 2
challengeType: 1
dashedName: step-2
---

# --description--

如你在先前的課程中所學，布林值是一種只有 `true` 和 `false` 值的資料型別。

truthy 值是在布林值背景關係中評估為 `true` 的值。truthy 值的範例有：

- 非空字串
- 非零數字
- 布林值 `true`

以及你稍後會遇到的其他值。

你可以使用 `Boolean()` 函式來檢查一個值的真值。例如，`Boolean("Hello World!")` 會傳回 `true`，因為 `"Hello World!"` 是真值。

現在，呼叫 `Boolean(truthyOrFalsy)` 並將結果記錄到主控台。

# --hints--

你應該將 `Boolean(truthyOrFalsy)` 記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const truthyOrFalsy = true;
--fcc-editable-region--

--fcc-editable-region--
```
