---
id: 690a3ce03013be44d116136e
title: 步驟 12
challengeType: 1
dashedName: step-12
---

# --description--

條件述語可以有一個 `else` 語句，當 `if` 條件為假時會執行程式碼。以下是一個 `if...else` 述語的範例：

```js
if (condition) {
  console.log("condition is truthy");
} else {
  console.log("condition is falsy");
}
```

為你的現有 `if` 述語添加一個 `else` 語句。在你的 `else` 語句主體內，將 `"Timmy is not old enough to drive."` 記錄到主控台。

# --hints--

你的第三個 `if` 述語應該有一個 `else` 語句。

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}\s+else\s+{/s);
```

你應該在你的 `else` 語句主體內將 `"Timmy is not old enough to drive."` 記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}\s+else\s+{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is not old enough to drive.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 18;

if (timmyAge >= 16) {
  console.log("Timmy is old enough to drive.");
--fcc-editable-region--
}
--fcc-editable-region--
```
