---
id: 690a57fef66aac39dffc61a2
title: 步驟 5
challengeType: 1
dashedName: step-5
---

# --description--

既然你知道什麼是真值（truthy）和假值（falsy），就從你的程式碼中移除你的變數宣告和 `console.log` 述語。

# --hints--

你不應該有 `truthyOrFalsy` 變數。

```js
assert.notMatch(__helpers.removeJSComments(code), /const\s+truthyOrFalsy\s*=/);
```

你不應該在你的程式碼中使用 `console.log(Boolean(truthyOrFalsy))`。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "";

console.log(Boolean(truthyOrFalsy));
--fcc-editable-region--
```
