---
id: 66f53dc2c5bd6a11d6c3282f
title: 비밀번호 생성기 앱 만들기
challengeType: 26
dashedName: lab-password-generator
---

# --description--

이 실습에서는 함수를 사용하여 무작위 비밀번호 생성기를 만드는 연습을 합니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 길이를 나타내는 매개변수를 받는 `generatePassword`라는 함수를 만들어야 합니다. 매개변수 이름은 자유롭게 정하세요.
2. 함수는 무작위로 생성된 비밀번호를 나타내는 문자열을 반환해야 합니다. 다음 문자열과 다양한 `Math` 메서드를 사용하여 무작위 문자가 포함된 새 문자열을 반환하세요: `ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()`.
3. `password`라는 변수를 정의하고, 원하는 비밀번호 길이를 나타내는 숫자 인수를 사용해 `generatePassword` 함수를 호출한 결과를 할당해야 합니다.
4. `console.log`가 있어야 하며, 메시지 `Generated password:`와 `password` 변수를 공백으로 구분해 연결한 단일 문자열을 로그로 출력해야 합니다.

# --hints--

매개변수가 있는 `generatePassword` 함수를 만들어야 합니다. 매개변수 이름은 자유롭게 정하세요.

```js
assert.isFunction(generatePassword);
assert.lengthOf(generatePassword, 1);
```

`generatePassword` 함수는 문자열을 반환해야 합니다.

```js
const result = generatePassword(5);
assert.isString(result);
```

`generatePassword` 함수는 올바른 길이의 새 문자열을 반환해야 합니다.

```js
const length = 8;
const password = generatePassword(length);
assert.lengthOf(password, length);
```

함수는 유효한 문자를 포함한 무작위로 생성된 비밀번호를 반환해야 합니다.

```js
const validChars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
const password = generatePassword(10);
for (let char of password) {
    assert.include(validChars, char);
}
```

함수는 하나 이상의 고유 문자를 포함하는 무작위 비밀번호를 반환해야 합니다.

```js
const passwordSetSize1 = new Set(generatePassword(12)).size;
const passwordSetSize2 = new Set(generatePassword(12)).size;
const passwordSetSize3 = new Set(generatePassword(12)).size;
assert.isAbove(passwordSetSize1 + passwordSetSize2 + passwordSetSize3, 3);
```

함수는 호출할 때마다 새 무작위 문자열을 반환해야 합니다.

```js
const password1 = generatePassword(10);
const password2 = generatePassword(10);
assert.notStrictEqual(password1, password2);
```

`password` 변수가 있어야 합니다.

```js
assert.isDefined(password);
```

`password` 변수는 문자열이어야 합니다.

```js
assert.isString(password);
```

숫자 인수를 사용해 `generatePassword` 함수를 호출하고 반환된 비밀번호를 `password` 변수에 저장해야 합니다.

```js
assert.isDefined(password);
assert.isString(password);
assert.match(password, /^[A-Za-z0-9!@#$%^&*()]+$/);
const length2 = 10;
const password3 = generatePassword(length2);
assert.lengthOf(password3, length2);
assert.equal(password3.length, generatePassword(length2).length);
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+password\s*=\s*generatePassword\(\d+\)\;?/);
```

`Generated password:`와 `password`를 공백 하나로 구분해 연결한 단일 문자열을 `+` 또는 템플릿 리터럴을 사용해 로그로 출력해야 합니다.

```js
const cleanCode = __helpers.removeJSComments(code);
const condition1 = /console\.log\(\s*["']Generated\s+password:\s*["']\s*\+\s*password\s*\);?/gm.test(cleanCode);
const condition2 = /console\.log\(\s*`Generated\s+password:\s*\$\{password\}`\s*\);?/gm.test(cleanCode);
assert.isTrue(condition1 || condition2);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const generatePassword = (length) => {
    const chars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
    let password = "";
    for (let i = 0; i < length; i++) {
        const randomIndex = Math.floor(Math.random() * chars.length);
        password += chars[randomIndex];
    }
    return password;
};

const password = generatePassword(12);
console.log(`Generated password: ${password}`);
```
