---
id: 686b9720ee1d032bd77a480a
title: تصحيح مدقق ISBN
challengeType: 27
dashedName: debug-an-isbn-validator
---

# --description--

رقم الكتاب المعياري الدولي (ISBN) هو معرف فريد يُعيَّن للكتب التجارية. يمكن أن يكون طوله 10 أو 13 رقمًا، والرقم الأخير هو رقم تحقق يُحسب من الأرقام الأخرى.

حاول Camperbot بناء مدقق ISBN خاص به. مع ذلك، ارتكب بعض الأخطاء على طول الطريق.

في هذا المختبر، ستصلح الشفرة الحالية وتجعلها تعمل بشكل صحيح.

السلوك المتوقع:
عندما يشغل المستخدم البرنامج، سيظهر له الموجه `Enter ISBN and length: `. يمكن للمستخدم إدخال رمز ISBN الذي يريد التحقق منه بصيغة `ISBN,length`. يجب ألا يحتوي رمز ISBN على شرطات، يليه طوله (`10` أو `13`)، مفصولًا بفاصلة.

أمثلة على المدخلات:
`1530051126,10` لرموز ISBN-10.
`9781530051120,13` لرموز ISBN-13.

كيفية إيجاد رقم التحقق:

- لا تحتاج إلى معرفة منطق الحساب التفصيلي في هذا المختبر. الدالتان `calculate_check_digit_10` و `calculate_check_digit_13` ستتوليان الحساب وترجعان رقم التحقق المتوقع كسلسلة نصية.
- رقم التحقق لرموز ISBN-10 يمكن أن يكون رقمًا من `0` إلى `9` أو الحرف الكبير `X`.
- رقم التحقق لرموز ISBN-13 يمكن أن يكون رقمًا من `0` إلى `9`.

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب عليك إصلاح `IndentationError` في الشفرة الحالية.
1. حتى لو لم يدخل المستخدم قيمة مفصولة بفاصلة، يجب على البرنامج التعامل مع `IndexError` دون تعطل.
1. عندما لا يدخل المستخدم قيمة مفصولة بفاصلة، يجب أن يرى الرسالة `Enter comma-separated values.` في وحدة التحكم، وينهي البرنامج.
1. حتى لو أدخل المستخدم قيمة غير رقمية للطول، يجب على البرنامج التعامل مع `ValueError` دون تعطل.
1. عندما يدخل المستخدم قيمة غير رقمية للطول، يجب أن يرى الرسالة `Length must be a number.` في وحدة التحكم، وينهي البرنامج.
1. يجب عليك إصلاح خطأ بمقدار واحد في دالة `validate_isbn`.
1. يجب عليك إصلاح `TypeError` في الشفرة الحالية التي تحدث عندما يدخل المستخدم رمز ISBN صالح.
1. يجب عليك إصلاح `IndexError` في الشفرة الحالية عندما يدخل المستخدم رمز ISBN صالح.
1. حتى لو أدخل المستخدم رمز ISBN غير صحيح يحتوي على حروف غير أرقام، يجب على البرنامج التعامل مع `ValueError` التي تحدث دون تعطل.
1. عندما يدخل المستخدم رمز ISBN غير صحيح يحتوي على حروف غير أرقام، يجب أن يرى الرسالة `Invalid character was found.` في وحدة التحكم.
1. عندما يدخل المستخدم `1530051126,10`، يجب أن يرى الرسالة `Valid ISBN Code.`
1. عندما يدخل المستخدم `9781530051120,13`، يجب أن يرى الرسالة `Valid ISBN Code.`

**مهم**: ستحتاج إلى إضافة تعليق على استدعاء `main()` في النطاق العام لكي تعمل الاختبارات بشكل صحيح.

عند إكمال المشروع، يجب أن يرى المستخدم الرسائل التالية حسب القيم التي يدخلها.

| شفرة ISBN | الطول | الرسالة | مثال الإدخال |
|-----------|--------|---------| ------------- |
| صالح | صالح | `Valid ISBN Code.` | `1530051126,10` |
| رقم غير صالح | صالح | `Invalid ISBN Code.` | `1530051125,10` |
| لا يطابق الطول المحدد أو ترك فارغًا | صالح | `ISBN-10 code should be 10 digits long.` أو `ISBN-13 code should be 13 digits long.`، حسب الطول الذي أدخله المستخدم. | `9781530051120,10` أو `1530051126,13` |
| يحتوي على أحرف غير رقمية (باستثناء رقم التحقق) | صالح | `Invalid character was found.` | `15-0051126,10` |
| أي | رقم غير صالح | `Length should be 10 or 13.` | `1530051126,9` |
| أي | يحتوي على أحرف غير رقمية أو ترك فارغًا | `Length must be a number.` | `1530051125,A` |
| غير مفصول بفاصلة | غير مفصول بفاصلة | `Enter comma-separated values.` | `1530051125` |

يمكنك استخدام القيم التالية لاختبار شفرتك يدويًا إذا أردت.

أمثلة على المدخلات لرموز ISBN-10 الصالحة:
`1530051126,10`
`9971502100,10`
`080442957X,10`

أمثلة على المدخلات لرموز ISBN-13 الصالحة:
`9781530051120,13`
`9781947172104,13`

# --hints--

يجب عليك إضافة تعليق على استدعاء دالة `main` للسماح لبقية الاختبارات بالعمل بشكل صحيح.

```js
({
  test: () => runPython('assert not _Node(_code).has_call("main()")')
})
```

يجب أن تحتوي على دالة `validate_isbn`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("validate_isbn")')
})
```

يجب أن تحتوي على دالة `calculate_check_digit_10`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_10")')
})
```

يجب أن تحتوي على دالة `calculate_check_digit_13`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_13")')
})
```

يجب أن تحتوي على دالة `main`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("main")')
})
```

عندما يدخل المستخدم قيمة غير مفصولة بفاصلة، يجب أن ترى الرسالة `Enter comma-separated values.` في وحدة التحكم.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

عندما يدخل المستخدم قيمة غير رقمية للطول، يجب أن ترى الرسالة `Length must be a number.` في وحدة التحكم.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,a"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

عندما يدخل المستخدم رمز ISBN غير صحيح يحتوي على حروف غير أرقام، يجب أن ترى الرسالة `Invalid character was found.` في وحدة التحكم.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "978IS3OOSII2O,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

عندما يدخل المستخدم `1530051126,10`، يجب أن ترى الرسالة `Valid ISBN Code.` في وحدة التحكم.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

عندما يدخل المستخدم `9781530051120,13`، يجب أن ترى الرسالة `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

عندما يدخل المستخدم `1530051125,10`، يجب أن ترى الرسالة `Invalid ISBN Code.`.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid ISBN Code." in out

`)
}})
```

عندما يدخل المستخدم `9781530051120,10`، يجب أن يرى الرسالة `ISBN-10 code should be 10 digits long.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-10 code should be 10 digits long." in out

`)
}})
```

عندما يدخل المستخدم `1530051126,13`، يجب أن يرى الرسالة `ISBN-13 code should be 13 digits long.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-13 code should be 13 digits long." in out

`)
}})
```

عندما يدخل المستخدم `15-0051126,10`، يجب أن يرى الرسالة `Invalid character was found.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "15-0051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

عندما يُدخل المستخدم `1530051126,9`، يجب أن يرى الرسالة `Length should be 10 or 13.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,9"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length should be 10 or 13." in out

`)
}})
```

عندما يدخل المستخدم `1530051125,A`، يجب أن يرى الرسالة `Length must be a number.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,A"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

عندما يدخل المستخدم `1530051125`، يجب أن ترى الرسالة `Enter comma-separated values.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

عندما يدخل المستخدم `9971502100,10`، يجب أن ترى الرسالة `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9971502100,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

عندما يدخل المستخدم `080442957X,10`، يجب أن ترى الرسالة `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "080442957X,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

عندما يدخل المستخدم `9781947172104,13`، يجب أن ترى الرسالة `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781947172104,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

# --seed--

## --seed-contents--

```py
def validate_isbn(isbn, length):
    if len(isbn, length) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return
    main_digits = isbn[0:length]
    given_check_digit = isbn[length]
    main_digits_list = [int(digit) for digit in main_digits]
    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)
    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')
def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11
    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10
    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]
    length = int(values[1])
    if length == 10 or length == 13:
    validate_isbn(isbn, length)
    else:
    print('Length should be 10 or 13.')

main()
```

# --solutions--

```py
def validate_isbn(isbn, length):
    if len(isbn) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return

    check_digit_index = length - 1

    main_digits = isbn[0:check_digit_index]
    given_check_digit = isbn[check_digit_index]

    try:
        main_digits_list = [int(digit) for digit in main_digits]
    except ValueError as e:
        print('Invalid character was found.')
        return

    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)

    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')


def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11

    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10

    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]

    try:
        length = int(values[1])
    except ValueError as e:
        print('Length must be a number.')
        return
    except IndexError as e:
        print('Enter comma-separated values.')
        return

    if length == 10 or length == 13:
        validate_isbn(isbn, length)
    else:
        print('Length should be 10 or 13.')


# main()
```
