---
id: 694a9ff65d4472dd451d8255
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Workshop üben Sie, wie Sie Booleans und bedingte Anweisungen in Python verwenden, indem Sie einen Rechner für die Buchung von Kinokarten erstellen.

Beginnen Sie damit, Variablen zu deklarieren, die Informationen über den Benutzer und die Filmvorstellung speichern.

Erstellen Sie eine Variable namens `base_price`, um den Grundpreis der Kinokarte zu speichern, und setzen Sie ihren Wert auf `15`. Erstellen Sie eine weitere Variable namens `age`, um das Alter des Benutzers zu speichern, und setzen Sie ihren Wert auf `21`.

# --hints--

Sie sollten eine Variable namens `base_price` deklarieren.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("base_price")`)) })
```

Sie sollten den ganzzahligen Wert `15` `base_price` zuweisen.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("base_price").is_equivalent("base_price = 15")
`)) })
```

Sie sollten eine Variable namens `age` deklarieren.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("age")`)) })
```

Sie sollten den ganzzahligen Wert `21` `age` zuweisen.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("age").is_equivalent("age = 21")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
