---
id: 694cc1b515976889afd4cded
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Jetzt müssen Sie einige String-Werte zum Kinoticket speichern.

Erstellen Sie eine Variable namens `seat_type`, um den Sitztyp zu speichern, den der Benutzer ausgewählt hat, und weisen Sie ihr den String `Gold` zu. Erstellen Sie eine weitere Variable namens `show_time`, um die Vorführzeit des Films zu speichern, und weisen Sie ihr den String `Evening` zu.

# --hints--

Sie sollten eine Variable namens `seat_type` deklarieren.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("seat_type")`)) })
```

Sie sollten den String `Gold` `seat_type` zuweisen. Denken Sie daran, den Wert entweder mit einfachen oder doppelten Anführungszeichen zu umgeben.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("seat_type").is_equivalent("seat_type = 'Gold'")
`)) })
```

Sie sollten eine Variable namens `show_time` deklarieren.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("show_time")`)) })
```

Sie sollten den String `Evening` `show_time` zuweisen. Denken Sie daran, den Wert entweder mit einfachen oder doppelten Anführungszeichen zu umgeben.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("show_time").is_equivalent("show_time = 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
--fcc-editable-region--

--fcc-editable-region--
```
