---
id: 6958e2377cacbc46a7f7383a
title: Schritt 13
challengeType: 20
dashedName: step-13
---

# --description--

Fügen Sie nun eine `else`-Klausel zu Ihrer `if`-Anweisung hinzu und geben Sie `No extra charges will be applied` innerhalb des `else`-Blocks aus. Dadurch wird die Nachricht nur ausgegeben, wenn die Bedingung für zusätzliche Gebühren falsch ist.

Verwenden Sie dann unterhalb der `else`-Klausel den `print()`-Aufruf, um eine Nachricht anzuzeigen, die `Extra charges:` gefolgt vom aktualisierten Wert von `extra_charges` zeigt, und überprüfen Sie die Ausgabe im Terminal.

# --hints--

Sie sollten eine `else`-Klausel haben.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3].find_conditions()[1].is_empty()`)) })
```

Sie sollten `No extra charges will be applied` innerhalb Ihrer `else`-Klausel ausgeben. Denken Sie daran, die Nachricht mit einfachen oder doppelten Anführungszeichen zu umgeben.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[1].has_call("print('No extra charges will be applied')")
`)) })
```

Sie sollten `print('Extra charges:', extra_charges)` unterhalb Ihrer `if...else`-Anweisung haben.

```js
({ test: () => (runPython(`
if_stmt = """if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')"""
print_call = "print('Extra charges:', extra_charges)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
--fcc-editable-region--

--fcc-editable-region--
```
