---
id: 6958dfa0c74b237e1b56ce04
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Der Benutzer erhält einen Mitgliederrabatt, wenn er Mitglied ist.

Erstellen Sie eine Variable namens `discount` und weisen Sie ihr den Wert `0` zu. Damit speichern Sie den Rabatt, den der Benutzer für das Kinoticket erhält.

# --hints--

Sie sollten eine Variable namens `discount` deklarieren.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("discount")`)) })
```

Sie sollten den ganzzahligen Wert `0` `discount` zuweisen.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("discount").is_equivalent("discount = 0")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

--fcc-editable-region--

--fcc-editable-region--
```
