---
id: 6958e59845777391b86d7331
title: Schritt 18
challengeType: 20
dashedName: step-18
---

# --description--

In Python kann eine `if`-Anweisung auch innerhalb des Körpers einer anderen `if`-Anweisung platziert werden. Dies nennt man eine verschachtelte `if`-Anweisung.

Eine verschachtelte `if`-Anweisung ermöglicht es Ihnen, eine zusätzliche Bedingung erst zu prüfen, nachdem die erste Bedingung bereits erfüllt wurde. Die innere `if`-Anweisung wird nur ausgeführt, wenn die äußere `if`-Bedingung wahr ist.

```py
if condition1:
    # Code to execute if condition1 is True
    if condition2:
        # Code to execute if both conditions are True
```

Jetzt berechnen Sie Servicegebühren basierend auf der Art des Sitzplatzes, den der Benutzer ausgewählt hat.

Innerhalb des Körpers der letzten `if`-Anweisung, unterhalb der `print('Ticket booking condition satisfied')`-Zeile, erstellen Sie eine Variable namens `service_charges` und setzen Sie diese auf `0`. Achten Sie darauf, Ihren Code um vier Leerzeichen einzurücken, damit er im Körper der äußeren `if`-Anweisung bleibt.

Erstellen Sie dann eine verschachtelte `if`-Anweisung, um zu prüfen, ob `seat_type` gleich `Premium` ist. Aktualisieren Sie innerhalb des Körpers der verschachtelten `if`-Anweisung den Wert von `service_charges` auf `5`.

# --hints--

Sie sollten eine Variable namens `service_charges` innerhalb des Körpers Ihrer letzten `if`-Anweisung deklarieren.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_variable("service_charges")
`)) })
```

Sie sollten den ganzzahligen Wert `0` an `service_charges` zuweisen.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 0")
`)) })
```

Sie sollten eine `if`-Anweisung innerhalb des Körpers Ihrer letzten `if`-Anweisung haben.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_ifs()[0]`)) })
```

Die Bedingung der neuen `if`-Anweisung sollte `seat_type == 'Premium'` sein.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("seat_type == 'Premium'")
`)) })
```

Sie sollten der Variablen `5` innerhalb Ihrer neuen `service_charges`-Anweisung den ganzzahligen Wert `if` zuweisen.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_ifs()[0].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 5")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

if age >= 21 or age >= 18 and (show_time != 'Evening' or is_member):
    print('Ticket booking condition satisfied')

--fcc-editable-region--
    
--fcc-editable-region--
else:
    print('Ticket booking failed due to restrictions')
```
