---
id: 6958e0ecc121b0ae3dd9c5e0
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Sie sollten auch den Fall behandeln, wenn der Benutzer keinen Anspruch auf einen Mitgliederrabatt hat.

Fügen Sie eine `else`-Klausel zu Ihrer `if`-Anweisung hinzu und geben Sie `User does not qualify for membership discount` im Körper von `else` aus.

Sie können eine Nachricht gefolgt von einer Variablen so ausgeben:

```py
print('Total:', total)
```

Sie möchten außerdem den aktualisierten Wert von `discount` anzeigen. Verwenden Sie unterhalb der `if...else`-Anweisung den `print()`-Aufruf, um eine Nachricht anzuzeigen, die `Discount:` gefolgt vom aktualisierten Wert von `discount` zeigt. Überprüfen Sie dann die Ausgabe im Terminal.

# --hints--

Sie sollten eine `else`-Klausel haben.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2].find_conditions()[1].is_empty()`)) })
```

Sie sollten `User does not qualify for membership discount` innerhalb Ihrer `else`-Klausel ausgeben.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[1].has_call("print('User does not qualify for membership discount')")
`)) })
```

Sie sollten `print('Discount:', discount)` unterhalb Ihrer `if...else`-Anweisung haben.

```js
({ test: () => (runPython(`
if_stmt = """if is_member:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')"""
print_call = "print('Discount:', discount)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
if is_member:
    discount = 3
    print('User qualifies for membership discount')
--fcc-editable-region--

--fcc-editable-region--
```
