---
id: 6958e1875e592b58e8852757
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Erstellen Sie nun eine Variable namens `extra_charges` und weisen Sie ihr den Wert `0` zu. Diese Variable stellt zusätzliche Gebühren dar, die an Wochenenden auf das Kinoticket angewendet werden.

Erstellen Sie anschließend eine `if`-Anweisung, um zu prüfen, ob `is_weekend` wahr ist. Innerhalb des Körpers der `if`-Anweisung aktualisieren Sie den Wert von `extra_charges` auf `2` und geben `Extra charges will be applied` im Terminal aus.

# --hints--

Sie sollten eine Variable namens `extra_charges` deklarieren.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("extra_charges")`)) })
```

Sie sollten den ganzzahligen Wert `0` `extra_charges` zuweisen.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("extra_charges").is_equivalent("extra_charges = 0")
`)) })
```

Sie sollten eine vierte `if`-Anweisung haben.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3]`)) })
```

Die Bedingung Ihrer vierten `if`-Anweisung sollte `is_weekend` sein.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_conditions()[0].is_equivalent("is_weekend")
`)) })
```

Sie sollten den ganzzahligen Wert `2` der Variablen `extra_charges` innerhalb Ihrer neuen `if`-Anweisung zuweisen.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[3].find_bodies()[0].find_variable("extra_charges") 
assert _var.is_equivalent("extra_charges = 2")
`)) })
```

Sie sollten `Extra charges will be applied` innerhalb Ihrer neuen `if`-Anweisung ausgeben. Denken Sie daran, die Nachricht mit einfachen oder doppelten Anführungszeichen zu umgeben.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[0].has_call("print('Extra charges will be applied')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

--fcc-editable-region--

--fcc-editable-region--
```
