---
id: 6958de0d773289a89515a6b8
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Jetzt prüfen Sie, ob der Benutzer basierend auf seinem Alter berechtigt ist, eine Abendvorstellung zu buchen.

Erstellen Sie eine `if`-Anweisung, um zu überprüfen, ob `age` größer oder gleich `21` ist. Innerhalb des Rumpfs der `if`-Anweisung geben Sie `User is eligible for Evening shows` aus.

# --hints--

Ihr Code sollte zwei `if`-Anweisungen enthalten.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1]`)) })
```

Die Bedingung Ihrer zweiten `if`-Anweisung sollte `age >= 21` sein.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[1].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

Sie sollten `User is eligible for Evening shows` innerhalb Ihrer neuen `if`-Anweisung ausgeben. Denken Sie daran, die Nachricht mit einfachen oder doppelten Anführungszeichen zu umschließen.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[0].has_call("print('User is eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

--fcc-editable-region--

--fcc-editable-region--
```
