---
id: 6960ba6e6c95ed4ad896c6f2
title: Schritt 17
challengeType: 20
dashedName: step-17
---

# --description--

Es gibt noch eine Ausnahme bei den Buchungsregeln. Benutzer zwischen `18` und `21` können Abendvorstellungen buchen, wenn sie Mitglieder sind.

Wenn in einer `if`-Anweisung mehrere logische Operatoren verwendet werden, werden Bedingungen, die mit `and` verbunden sind, vor Bedingungen ausgewertet, die mit `or` verbunden sind. Runde Klammern `()` werden in Python verwendet, um Bedingungen zu gruppieren und die Reihenfolge ihrer Auswertung zu steuern.

```py
if condition1 and (condition2 or condition3):
    # Code to execute if conditions evaluate to True
else:
    # Code to execute if conditions evaluate to False
```

Aktualisieren Sie die Bedingung der `if age >= 21 or age >= 18 and show_time != 'Evening':`-Zeile, um eine weitere Bedingung mit dem `or`-Operator hinzuzufügen, die überprüft, ob `is_member` wahr ist. Verwenden Sie runde Klammern `()`, um die `show_time != 'Evening'`- und `is_member`-Bedingungen wie im obigen Beispiel zusammenzufassen.

# --hints--

Die letzte Bedingung Ihrer `if`-Anweisung sollte `age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)` sein.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--
if age >= 21 or age >= 18 and show_time != 'Evening':
--fcc-editable-region--
    print('Ticket booking condition satisfied')
else:
    print('Ticket booking failed due to restrictions')
```
