---
id: 6960a5165f7efe52a0c0aa39
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

In Python wird der `and`-Operator verwendet, um zu prüfen, ob mehrere Bedingungen wahr sind. So können Sie ihn verwenden, um zwei Bedingungen in einer `if`-Anweisung zu kombinieren:

```py
if condition1 and condition2:
    # Code to execute if all conditions are True
```

Der Mitgliederrabatt soll nur für Mitglieder gelten, wenn ihr `age` größer oder gleich `21` ist.

Aktualisieren Sie die Bedingung der `if is_member:`-Zeile, indem Sie den `and`-Operator verwenden, um die bestehende Bedingung mit einer weiteren Bedingung zu kombinieren, die prüft, ob `age` größer oder gleich `21` ist.

# --hints--

Die dritte Bedingung Ihrer `if`-Anweisung sollte `is_member and age >= 21` sein.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[2].find_conditions()[0] 
assert _cond.is_equivalent("is_member and age >= 21") or _cond.is_equivalent("is_member and 21 <= age")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--
if is_member:
--fcc-editable-region--
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
