---
id: 696099abb6acc539b27055a1
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Im vorherigen Schritt haben Sie eine Bedingung mit einer `if`-Anweisung überprüft. Die `else`-Klausel wird verwendet, um den Fall zu behandeln, wenn die Bedingung falsch ist. Hier ist die Syntax einer `if...else`-Anweisung:

```py
if condition:
    # Code to execute if condition is True
else:
    # Code to execute if condition is False
```

Fügen Sie nun eine `else`-Klausel zu Ihrer `if`-Anweisung hinzu und geben Sie `User is not eligible for Evening shows` im Körper der `else` aus. Dies wird nur ausgegeben, wenn die Bedingung in der `if`-Anweisung falsch ist.

# --hints--

Sie sollten eine `else`-Klausel haben.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1].find_conditions()[1].is_empty()`)) })
```

Sie sollten `User is not eligible for Evening shows` innerhalb Ihrer `else`-Klausel ausgeben. Denken Sie daran, die Nachricht mit einfachen oder doppelten Anführungszeichen zu umgeben.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[1].has_call("print('User is not eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
--fcc-editable-region--

--fcc-editable-region--
```
