---
id: 6960b6a6ea06e27dcc7eef37
title: Schritt 15
challengeType: 20
dashedName: step-15
---

# --description--

Jetzt überprüfen Sie, ob der Benutzer die Bedingungen erfüllt, um ein Kinoticket zu buchen. Benutzer mit einem Alter von `21` oder älter können immer ohne Einschränkungen Tickets buchen.

Erstellen Sie eine `if`-Anweisung, um zu prüfen, ob `age` größer oder gleich `21` ist. Geben Sie im Rumpf der `if`-Anweisung `Ticket booking condition satisfied` im Terminal aus.

Fügen Sie dann eine `else`-Klausel zu Ihrer `if`-Anweisung hinzu und geben Sie `Ticket booking failed due to restrictions` im Rumpf der `else` aus.

# --hints--

Sie sollten eine fünfte `if`-Anweisung haben.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4]`)) })
```

Die Bedingung Ihrer fünften `if`-Anweisung sollte `age >= 21` sein.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

Sie sollten `Ticket booking condition satisfied` in Ihrer neuen `if`-Anweisung ausgeben. Denken Sie daran, die Nachricht mit einfachen oder doppelten Anführungszeichen zu umgeben.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_call("print('Ticket booking condition satisfied')")
`)) })
```

Ihre neue `if`-Anweisung sollte eine `else`-Klausel haben.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_conditions()[1] == _Node()`)) })
```

Sie sollten `Ticket booking failed due to restrictions` in Ihrer `else`-Klausel ausgeben. Denken Sie daran, die Nachricht mit einfachen oder doppelten Anführungszeichen zu umgeben.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[1].has_call("print('Ticket booking failed due to restrictions')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--

--fcc-editable-region--
```
