---
id: 696099c1ac72a4e585de1d5b
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Manche Informationen können nur wahr oder falsch sein. Wie Sie in vorherigen Lektionen gelernt haben, lässt sich das mit Boolean-Werten darstellen.

Erstellen Sie eine Variable namens `is_member`, die angibt, ob der Benutzer Mitglied ist, und weisen Sie ihr den Wert `True` zu.

Erstellen Sie unterhalb der Variable `is_member` eine weitere Variable namens `is_weekend`, die angibt, ob die Filmvorführung am Wochenende stattfindet, und weisen Sie ihr den Wert `False` zu. Umgeben Sie den Wert nicht mit Anführungszeichen.

# --hints--

Sie sollten eine Variable namens `is_member` deklarieren.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_member")`)) })
```

Sie sollten den Boolean-Wert `True` der Variable `is_member` zuweisen. Umgeben Sie den Wert nicht mit Anführungszeichen.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = True")
`)) })
```

Sie sollten eine Variable namens `is_weekend` deklarieren.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_weekend")`)) })
```

Sie sollten den Boolean-Wert `False` der Variable `is_weekend` zuweisen. Umgeben Sie den Wert nicht mit Anführungszeichen.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_weekend").is_equivalent("is_weekend = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--

--fcc-editable-region--
```
