---
id: 696cb889fc4fdeb337e52864
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

In Python müssen `if`-Bedingungen Werte nicht explizit mit `True` oder `False` vergleichen. Stattdessen verlassen sie sich auf die truthiness von Werten. Truthy-Werte werden in einem Boolean-Kontext, wie der Bedingung einer `True`-Anweisung, als `if` ausgewertet. Dazu gehören alle Zahlen ungleich null, nicht-leere Strings und der Boolean-Wert `True`.

```py
if value:
   print('value is truthy')
```

Erstellen Sie eine `if`-Anweisung, um zu prüfen, ob `is_member` truthy ist. Innerhalb des Körpers der `if`-Anweisung aktualisieren Sie den Wert von `discount` auf `3` und geben `User qualifies for membership discount` im Terminal aus.

# --hints--

Ihr Code sollte drei `if`-Anweisungen enthalten.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2]`)) })
```

Die Bedingung Ihrer dritten `if`-Anweisung sollte `is_member` sein.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_conditions()[0].is_equivalent("is_member")
`)) })
```

Sie sollten der Variablen `3` innerhalb Ihrer neuen `discount`-Anweisung den ganzzahligen Wert `if` zuweisen.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].find_variable("discount").is_equivalent("discount = 3")
`)) })
```

Sie sollten `User qualifies for membership discount` innerhalb Ihrer neuen `if`-Anweisung ausgeben. Denken Sie daran, die Nachricht mit einfachen oder doppelten Anführungszeichen zu umgeben.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].has_call("print('User qualifies for membership discount')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--

--fcc-editable-region--
```
