---
id: 65c4f37ba6cbcfdc77e24165
title: الخطوة 17
challengeType: 20
dashedName: step-17
---

# --description--

سيتم استدعاء طريقة `insert` من قبل المستخدم. بالإضافة إلى معلمة `self`، ستحتاج أيضًا إلى معلمة `key`. ستكون هذه المعلمة هي قيمة المفتاح التي تُدرج في شجرة البحث الثنائي.

أضف `key` كالمعلمة الثانية لتعريف الدالة.

# --hints--

يجب أن تحتوي طريقة `insert` على المعلمة `self`.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,/) });
```

يجب أن تحتوي طريقة `insert` على المعلمة `key`.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

--fcc-editable-region--            
    def insert(self):
        pass

--fcc-editable-region--

```
