---
id: 65c4eb814cc977c95cd7df0e
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذا المشروع، ستنشئ شجرة البحث الثنائي (BST). شجرة البحث الثنائي هي هيكل بيانات تحتوي كل عقدة فيها على طفلين كحد أقصى، حيث يحتوي الطفل الأيسر على قيم أقل من عقدة الأصل، والطفل الأيمن يحتوي على قيم أكبر من عقدة الأصل، مما يسمح بعمليات بحث وفرز فعالة.

هذا هو شكل شجرة البحث الثنائي:

<img alt="a binary search tree" src="https://cdn.freecodecamp.org/curriculum/python/bst-example.png" style="background-color: white; height:500px; width:500px; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

ابدأ بتعريف صنف `TreeNode` فارغ. يمثل صنف `TreeNode` عقدة في شجرة البحث الثنائي. استخدم الكلمة المفتاحية `pass` لملء جسم الصنف وتجنب حدوث خطأ.

# --hints--

يجب أن تستخدم الكلمة المفتاحية `class` لإعلان صنف فارغ باسم `TreeNode`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(TreeNode)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
