---
id: 65c4f300da28d8da361bfa93
title: الخطوة 14
challengeType: 20
dashedName: step-14
---

# --description--

داخل بند `elif`، استبدل `pass` باستدعاء الطريقة `_insert` مع تمرير الابن الأيمن للعقدة الحالية كالمعلمة الأولى و`key` كالمعلمة الثانية. عيّن النتيجة إلى الخاصية `right` للعقدة الحالية.

# --hints--

يجب إزالة الكلمة المفتاحية `pass` من كتلة `elif`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

يجب استدعاء الطريقة `self._insert()` مع تمرير `node.right` و`key` كمعلمات.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```

يجب تعيين نتيجة استدعاء `self._insert()` إلى الخاصية اليمنى (`node.right`) للعقدة الحالية.

```js
({ test: () => assert.match(code, /node\.right\s*=\s*self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```


# --seed--

## --seed-contents--


```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            pass

--fcc-editable-region--

```
