---
id: 65c4f0b82db314d3fcc0b8e1
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

داخل صنف `BinarySearchTree`، استبدل `pass` بطريقة `__init__` وأضف معلمة `self` إلى هذه الطريقة.

# --hints--

يجب إزالة الكلمة المفتاحية `pass` من صنف `BinarySearchTree`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});

```


يجب تعريف طريقة `__init__` وإضافة معلمة `self` إلى هذه الطريقة. تذكر استخدام الكلمة المفتاحية `pass` داخل جسم الطريقة.

```js
({ test: () => assert.match(code, /^\s{4}def\s+__init__\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
        
--fcc-editable-region--
class BinarySearchTree:
    pass
--fcc-editable-region--
```
