---
id: 65c4f2851645e8d84d14f5d1
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

الآن عليك أن تتنقل في الشجرة بشكل متكرر وتُدرج القيم باستخدام المبدأ الخاص بالأشجار الثنائية:

- القيم الأصغر من المفتاح توضع في الشجرة الفرعية اليسرى
- القيم الأكبر من المفتاح توضع في الشجرة الفرعية اليمنى

بعد جملة الشرط الموجودة لديك، اكتب جملة `if` أخرى للتحقق مما إذا كانت قيمة `key` أقل من `node.key`.

# --hints--

يجب أن تكتب جملة if أخرى للتحقق مما إذا كان `key` أقل من `node.key`. تذكر استخدام الكلمة المفتاحية `pass`.

```js
({ test: () => assert(runPython(` _Node(_code).find_class("BinarySearchTree").find_function("_insert").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

--fcc-editable-region--
```

