---
id: 65c4f2b7178afed8e88f782f
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

إذا كانت نتيجة `key` < `node.key` تساوي `True`، فيجب وضع العقدة الجديدة في الشجرة الفرعية اليسرى.

احذف `pass` واستدعِ بطريقة متكررة الدالة `_insert` مع تمرير الابن الأيسر كالمعلمة الأولى و`key` كالمعلمة الثانية. عيّن النتيجة إلى الخاصية `left` للعقدة الحالية.

# --hints--

يجب إزالة الكلمة المفتاحية `pass` من كتلة `if`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_insert").find_if("key < node.key").find_bodies()[0].has_pass()`
      )
    );
  },
});
```

يجب استدعاء الطريقة `self._insert` مع تمرير `node.left` و`key` كمعلمات.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

يجب تعيين نتيجة استدعائك لـ `self._insert()` إلى الخاصية `left` للعقدة الحالية.

```js
({ test: () => assert.match(code, /node\.left\s*=\s*self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--    

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            pass
--fcc-editable-region--    

```
