---
id: 65c4f0d32ca17ad4b1636b0e
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

داخل الطريقة `__init__`، احذف `pass` وابدأ تهيئة `root` بالقيمة `None`.

تمثل الخاصية `root` العقدة الجذرية لشجرة البحث الثنائي. بما أن هذه هي دالة البناء عند إنشاء كائن جديد من `BinarySearchTree`، تبدأ الشجرة فارغة، لذا تُعيّن الخاصية `root` إلى `None`.

# --hints--

يجب حذف جملة `pass` من الطريقة `__init__`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

يجب تهيئة الخاصية `root` إلى `None` باستخدام `self.root`.

```js
({ test: () => assert.match(code, /^\s{8}self\.root\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        pass
--fcc-editable-region--        
```
