---
id: 65c4f195de7e2ad5932be717
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

بعد ذلك، تحتاج إلى تعريف آلية لإدخال العقد في الشجرة. لهذا، عليك تعريف طريقة `_insert`، وهي دالة مساعدة ستُستخدم لاحقًا بواسطة الطريقة الفعلية `insert`.

هذه الطريقة متكررة، بمعنى أنها تستدعي نفسها لتجتاز الشجرة حتى تجد الموقع المناسب للعقدة الجديدة.

عرّف طريقة `_insert` بالمعلمات `self` و`node` و`key`.

# --hints--

يجب أن تعرف طريقة `_insert` داخل صنف `BinarySearchTree`. تذكر استخدام `pass`.

```js
({ test: () => assert.match(code, /def\s+_insert\s*\([^(]*\)\s*:/m) })
```

يجب أن تأخذ طريقة `_insert` ثلاث معلمات: `self` و`node` و`key`.

```js
({ test: () => assert.match(code, /def\s+_insert\s*\(\s*self\s*,\s*node\s*,\s*key\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--
```
