---
id: 65c4f09e074dd8d37830ea00
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

أنشئ صنفًا فارغًا آخر باسم `BinarySearchTree` يمثل شجرة البحث الثنائي.

# --hints--

يجب أن يكون لديك صنف باسم `BinarySearchTree`. تذكر استخدام الكلمة المفتاحية `pass`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(BinarySearchTree)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--
```
