---
id: 65c4f013851cefd1a4fe4c96
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

تأخذ الطريقة `__init__` معلمتين: `self` (التي تمثل نسخة الصنف الجاري إنشاؤها) و`key` (القيمة التي ستُخزن في العقدة). أضف هاتين المعلمتين إلى الطريقة `__init__()`.

# --hints--

يجب إضافة المعلمة `self` إلى الطريقة.

```js
({ test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*/) });
```

يجب إضافة المعلمة `key` إلى الطريقة.

```js
({
  test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__():
        pass
--fcc-editable-region--
```
