---
id: 65c4f22498d22ed775ef8efb
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

الآن عليك التحقق مما إذا كانت معلمة `node` تساوي `None`. إذا كانت كذلك، فهذا يعني أن الطريقة وصلت إلى عقدة ورقية أو مكان فارغ في الشجرة حيث يجب إدخال العقدة الجديدة.

داخل جسم طريقة `_insert`، استبدل `pass` بعبارة `if` تتحقق مما إذا كانت `node is None`.

داخل كتلة `if` الجديدة، أرجع `TreeNode(key)` لإنشاء نسخة جديدة من `TreeNode` بالمفتاح المقدم. ستصبح هذه العقدة الورقية الجديدة، مما يُدخل المفتاح في الشجرة فعليًا.

# --hints--

يجب ألا تحتوي طريقة `_insert` على `pass`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_insert").has_pass()`)) })

```

يجب أن تكتب عبارة `if` للتحقق مما إذا كانت `node is None`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^\s{8}if\s+node\s+is\s+None\s*:/m));
  }
})
```

يجب أن تُرجع `TreeNode(key)` من كتلة `if`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^(\s{8})if\s+node\s+is\s+None\s*:\s*^\1\s{4}return\s+TreeNode\s*\(\s*key\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        pass
--fcc-editable-region--
```
