---
id: 65c4f07da204f4d2f325fbd3
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

داخل الطريقة `__init__`، ابدأ بتهيئة الخاصيتين `left` و `right` للعقدة إلى `None`. ذلك لأن العقدة عند إنشائها لأول مرة لا تحتوي على أبناء يسار أو يمين. تذكر استخدام الكلمة المفتاحية `self`.

# --hints--

يجب تعيين `None` إلى الخاصية `left` للعقدة باستخدام `self.left`.

```js
({ test: () => assert.match(code, /self\.left\s*=\s*None/m) })
```

يجب تعيين `None` إلى الخاصية `right` للعقدة باستخدام `self.right`.

```js
({ test: () => assert.match(code, /self\.right\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:    
    def __init__(self, key):
        self.key = key
--fcc-editable-region--
```
