---
id: 65c4ef8463f869d0eaf87c67
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

داخل صنف `TreeNode`، استبدل `pass` بطريقة `__init__` حتى تتمكن من تهيئة خصائص الكائن. لا تضف أي معلمات في الوقت الحالي.

# --hints--

يجب إزالة الكلمة المفتاحية `pass` من صنف `TreeNode` ونقلها داخل طريقة `__init__`.

```js
({
  test: () => {
    const pyClassStr = runPython(`str(_Node(_code).find_class("TreeNode"))`);
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});
```

يجب تعريف طريقة `__init__` داخل صنف `TreeNode`. تذكر استخدام الكلمة المفتاحية `pass` بداخلها.

```js
({
  test: () => {
    assert(
      runPython(`_Node(_code).find_class("TreeNode").has_function("__init__")`)
    );
  },
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
--fcc-editable-region--
```
