---
id: 65c4f02dbd3557d23f12db62
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

داخل الطريقة `__init__`، احذف `pass` وقم بتعيين قيمة المعلمة `key` إلى الخاصية `key` للعقدة باستخدام `self.key`.

هذا يعني أن الخاصية `key` في مثيل `TreeNode` ستُعيّن إلى القيمة التي تم تمريرها أثناء إنشاء الكائن.

# --hints--

يجب حذف العبارة `pass` من الطريقة `__init__`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("TreeNode").find_function("__init__").has_pass()`
      )
    );
  },
});

```

يجب تعيين قيمة المعلمة `key` إلى الخاصية `key` للعقدة باستخدام `self.key`.

```js
({ test: () => assert.match(code, /^\s{8}self\.key\s*=\s*key/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__(self, key):
        pass
--fcc-editable-region--
```
