---
id: 65c4f33bc3c3d8db7f732787
title: الخطوة 16
challengeType: 20
dashedName: step-16
---

# --description--

الآن، لتنفيذ الإدخال الفعلي، عرّف طريقة `insert` فارغة داخل صنف `BinarySearchTree` وأعطها معلمة `self`.

# --hints--

يجب أن تعرّف طريقة `insert` بمعلمة `self` داخل صنف `BinarySearchTree`. تذكر كلمة المفتاح `pass`.

```js
({ test: () => assert.match(code, /^\s{4}def\s+insert\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node
--fcc-editable-region--    

--fcc-editable-region--

```
