---
id: 65c646d4148ae3b2d1cbcac4
title: الخطوة 30
challengeType: 20
dashedName: step-30
---

# --description--

لاحظ أن بحثك يُرجع شيئًا مثل `80: <__main__.TreeNode object at 0x108b3e0>`. هذا هو التمثيل الافتراضي كسلسلة نصية عند طباعة نسخة من صنف.

لتغيير ذلك لطباعة قيمة مفيدة، عرّف طريقة أخرى باسم `__str__` في صنف `TreeNode`. تأخذ معلمة واحدة هي `self`.

بعد تعريف `__str__` ستظهر لك استثناء في وحدة التحكم لأن طريقة `__str__` لا تُرجع شيئًا بعد. ستعمل على جسم الطريقة في الخطوة التالية.

# --hints--

يجب أن تعرف طريقة `__str__` تأخذ معلمة واحدة هي `self`. تذكر استخدام `pass`.

```js
assert.match(code, /^(\s+)def\s+__init__.+?^\1def\s+__str__\(\s*self\s*\)\s*:\s*\n^\1\1pass/ms)
```


# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

print('Search for 80:', bst.search(80))

```
