---
id: 65c645b838e7deb080fc25e0
title: الخطوة 29
challengeType: 20
dashedName: step-29
---

# --description--

الآن، تم إدخال العقد. لتتأكد من إدخالها بشكل صحيح، يمكنك البحث عن عقدة في الشجرة.

خارج حلقة for، ابحث عن العقدة `80` في مثيل `bst` وأضفها إلى استدعاء `print`. يجب أن تكون المعلمة الأولى في دالة `print` هي `'Search for 80:'`.

# --hints--

يجب أن تطبع نتيجة استدعاء `bst.search(80)` ويجب أن تحتوي جملة `print` على المعلمة الأولى `'Search for 80:'`.

```js
assert.match(code, /^print\(('|")Search for 80:('|"),\s*bst\.search\(80\)/gm);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

--fcc-editable-region--
```
