---
id: 65c63a4c4da62e9ae18e321a
title: الخطوة 19
challengeType: 20
dashedName: step-19
---

# --description--

حان الوقت للعمل على وظيفة البحث. تمامًا كما أنشأت طريقة مساعدة `_insert` للطريقة `insert`، تحتاج إلى إنشاء طريقة مساعدة `_search` للطريقة `search`.

عرّف الطريقة `_search` بثلاث معلمات، وهي `self`، و`node`، و`key`.

# --hints--

يجب إعلان طريقة `_search`. لاحظ الشرطة السفلية في البداية وتذكر استخدام `pass`.

```js
({ test: () => assert.match(code, /def\s+_search/) });
```

يجب أن تأخذ الطريقة `_search` ثلاث معلمات: `self`، و`node`، و`key`.

```js
({ test: () => assert.match(code, /def\s+_search\(\s*self\s*,\s*node\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
--fcc-editable-region--

--fcc-editable-region--
```
