---
id: 65c63fd9b387f0a6c5bf6a72
title: الخطوة 24
challengeType: 20
dashedName: step-24
---

# --description--

داخل طريقة `search`، احذف `pass` واستدعِ طريقة المساعدة `_search` مع المعلمات التالية.

- `self.root`: هذه هي جذر شجرة البحث الثنائي. يبدأ البحث من الجذر.
- `key`: هذه هي القيمة التي يريد المستخدم إيجادها في شجرة البحث الثنائي.

داخليًا، تُفوّض طريقة `search` منطق البحث الفعلي إلى طريقة المساعدة `_search` التي تُنفذ البحث التكراري الفعلي داخل شجرة البحث الثنائي.

# --hints--

يجب حذف الكلمة المفتاحية `pass` من طريقة `search`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("search").has_pass()`
      )
    );
  },
});

```

يجب استدعاء طريقة `_search` داخل طريقة `search`.

```js
({ test: () => assert(runPython(`
  str(_Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0]).startswith("self._search(")
`)) });
```

يجب استدعاء طريقة `_search` مع تمرير `self.root` و `key` كمعلمات.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0].is_equivalent("self._search(self.root, key)")
`)) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        pass
--fcc-editable-region--
```
