---
id: 65c63df529bd15a24c187c62
title: الخطوة 21
challengeType: 20
dashedName: step-21
---

# --description--

اكتب جملة `if` أخرى تتحقق مما إذا كانت المفتاح الهدف أقل من مفتاح العقدة الحالية.

داخل كتلة `if`، أعد نتيجة استدعاء الطريقة `_search` مع الابن الأيسر للعقدة الحالية و`key` كمعلمات.

# --hints--

يجب أن تكتب جملة `if` أخرى للتحقق مما إذا كان `key` < `node.key`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_search").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

يجب أن تُعيد `self._search(node.left, key)` من كتلة `if` الجديدة.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_search").find_ifs()[1].find_bodies()[0].is_equivalent("return self._search(node.left, key)")`)) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node

--fcc-editable-region--
```
