---
id: 65c63da2ed6769a10e141341
title: الخطوة 20
challengeType: 20
dashedName: step-20
---

# --description--

الآن ستعرّف حالة أساسية للبحث التكراري. احذف `pass` الحالي واكتب جملة `if` تتحقق من شرطين:

- إذا كانت `node` تساوي `None`: هذا يشير إلى أن البحث وصل إلى نهاية فرع دون العثور على المفتاح.
- إذا كانت `node.key == key`: هذا يعني أن المفتاح وُجد في العقدة الحالية.

ادمج الشرطين باستخدام عامل `or` وأرجع العقدة الحالية داخل كتلة `if`.

# --hints--

يجب حذف الكلمة المفتاحية `pass` من الطريقة `_search`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_search").has_pass()`
      )
    );
  },
});

```

يجب كتابة جملة `if` تتحقق مما إذا كانت `node` تساوي `None` أو إذا كانت `node.key` تساوي `key`.

```js
({ test: () => assert.match(code, /^\s{8}if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:/m) })
```

يجب إرجاع `node` من داخل كتلة `if`.

```js
({ test: () => assert.match(code, /^(\s{8})if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:\s*^\1\s{4}return\s+node/m) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        pass

--fcc-editable-region--
```
