---
id: 65c63e6962159fa372ecdbec
title: الخطوة 22
challengeType: 20
dashedName: step-22
---

# --description--

إذا لم تكن جملة `if` الثانية صحيحة `True`، فهذا يعني أن المفتاح الهدف أكبر من أو يساوي مفتاح العقدة الحالية.

في شجرة البحث الثنائي، إذا كان المفتاح الهدف أكبر من مفتاح العقدة الحالية، يستمر البحث في الشجرة الفرعية اليمنى.

بعد كتلة `if`، أعد نتيجة استدعاء الطريقة `_search` مع الابن الأيمن للعقدة الحالية والمفتاح كمعلمات.

# --hints--

يجب أن تستدعي الطريقة `_search` بشكل متكرر باستخدام `self._search()`.

```js
const after_split = code.split('return self._search(node.left, key)')[1];
assert.match(after_split, /self\._search/);
```

يجب أن تمرر معلمتين إلى استدعاء الطريقة `_search`: `node.right` و `key`.

```js
({ test: () => assert.match(code, /self\._search\(\s*node\.right\s*,\s*key\s*\)/) });
```

يجب أن تعيد نتيجة استدعاء الطريقة `_search`.

```js
({
  test: () => assert.match(code, /return\s+self\._search\(\s*node\.right\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)

--fcc-editable-region--
```
