---
id: 65c6409418806da8d0636ffc
title: الخطوة 26
challengeType: 20
dashedName: step-26
---

# --description--

وظائف `insert` و `search` مكتملة، حان الوقت لاختبارها.

لاحظ أنه في هذه المرحلة، العقد غير مرتبة ومُدرجة فقط. ستعمل على الترتيب باستخدام الترحيل inorder لاحقًا.

يمكنك إنشاء نسخة من صنف في Python بهذه الطريقة:

```python
object_name = ClassName()
```

خارج تعريفات الأصناف، أنشئ نسخة من صنف `BinarySearchTree` وقم بتعيينها إلى المتغير `bst`.


# --hints--

يجب إنشاء نسخة من صنف `BinarySearchTree`.

```js
assert.match(code, /BinarySearchTree\(\s*\)/);
```

يجب تعيين النسخة الجديدة من `BinarySearchTree` إلى المتغير `bst`.

```js
assert.match(code, /bst\s*=\s*BinarySearchTree\(\s*\)/);
``` 

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--


--fcc-editable-region--
```
