---
id: 65ca07dd6fa8840491b7a5cd
title: الخطوة 41
challengeType: 20
dashedName: step-41
---

# --description--

إذا لم يتحقق أي من الشرطين السابقين، فهذا يعني أن العقدة لها أبناء على اليسار واليمين.

لاختيار الخليفة، تحتاج إلى إيجاد الحد الأدنى للقيمة في الشجرة الفرعية اليمنى.
أصغر قيمة ستكون الخليفة بالترتيب للعقدة الحالية.

لإيجاد أصغر قيمة، أنشئ دالة مساعدة `_min_value` تأخذ معلمتين: `self` و `node`.

# --hints--

يجب تعريف طريقة `_min_value` مع `self` و `node` كمعلمتين. تذكر استخدام الكلمة المفتاحية `pass`.

```js
({ test: () => assert.match(code, /def\s+_min_value\(\s*self\s*,\s*node\s*\)\s*:/) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left   
--fcc-editable-region--
    
--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)
    
# print('Search for 80:', bst.search(80))


```
