---
id: 65cb45d130c97cb459439fac
title: الخطوة 42
challengeType: 20
dashedName: step-42
---

# --description--

للعثور على أصغر قيمة في الشجرة الفرعية اليمنى، تحتاج إلى التكرار خلال الأبناء اليساريين للعقدة المعطاة حتى تصل إلى العقدة اليسارية الأبعد (الأصغر) في الشجرة الفرعية.

للقيام بذلك، اكتب حلقة `while` تعمل عندما تكون `node.left is not None` وانقل `pass` داخل كتلة الـ `while`. هذا الشرط يتحقق مما إذا كان هناك ابن يساري. طالما يوجد ابن يساري، تستمر الحلقة ويوجد قيمة أصغر يمكن العثور عليها.

# --hints--

يجب إزالة الكلمة المفتاحية `pass` من الطريقة `_min_value`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_min_value").has_pass()`
      )
    );
  },
});
```

يجب استخدام الشرط `node.left is not None` في حلقة الـ `while`.

```js
({ test: () => assert.match(code, /while\s+node\.left\s+is\s+not\s+None\s*:/) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left  

--fcc-editable-region--
    def _min_value(self, node):
        pass

--fcc-editable-region--   

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
