---
id: 65c4f37ba6cbcfdc77e24165
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

`insert` 方法将被用户调用。 除了 `self` 参数之外，它还需要 `key` 参数。 此参数将是插入二叉搜索树的关键值。

将 `key` 作为第二个参数添加到函数定义中。

# --hints--

`insert` 方法应该包含 `self` 参数。

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,/) });
```

`insert` 方法应该包含 `key` 参数。

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

--fcc-editable-region--            
    def insert(self):
        pass

--fcc-editable-region--

```
