---
id: 65c4f300da28d8da361bfa93
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

在 `elif` 子句中，将 `pass` 替换为对 `_insert` 方法的调用，其中当前节点的右子节点作为第一个参数，`key` 作为第二个参数。 将结果赋值给当前节点的 `right` 属性。

# --hints--

你应该从 `elif` 块中删除 `pass` 关键字。

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

你应该调用 `self._insert()` 方法并传递 `node.right` 和 `key` 作为参数。

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```

你应该将 `self._insert()` 调用的结果分配给当前节点的正确属性 (`node.right`)。

```js
({ test: () => assert.match(code, /node\.right\s*=\s*self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```


# --seed--

## --seed-contents--


```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            pass

--fcc-editable-region--

```
