---
id: 65c4f195de7e2ad5932be717
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

接下来，你需要定义一种在树中插入节点的机制。 为此，你需要定义一个 `_insert` 方法，它是一个辅助函数，稍后将被实际的 `insert` 方法使用。 

此方法是递归的，这意味着它会调用自身来遍历树，直到找到新节点的适当位置。

定义一个 `_insert` 方法，其参数为 `self`、`node` 和 `key`。

# --hints--

你应该在 `BinarySearchTree` 类中定义一个 `_insert` 方法。 记得使用 `pass`。

```js
({ test: () => assert.match(code, /def\s+_insert\s*\([^(]*\)\s*:/m) })
```

你的 `_insert` 方法应该采用三个参数：`self`、`node` 和 `key`。

```js
({ test: () => assert.match(code, /def\s+_insert\s*\(\s*self\s*,\s*node\s*,\s*key\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--
```
