---
id: 65c4f2b7178afed8e88f782f
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

如果 `key` < `node.key` 计算结果为 `True`，那么新建节点应放置在左子树中。

删除 `pass` 并以左孩子作为第一个参数、`key` 作为第二个参数递归调用 `_insert` 方法。 将结果赋值给当前节点的 `left` 属性。

# --hints--

你应该从 `if` 块中删除 `pass` 关键字。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_insert").find_if("key < node.key").find_bodies()[0].has_pass()`
      )
    );
  },
});
```

你应该调用 `self._insert` 方法并传递 `node.left` 和 `key` 作为参数。

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

你应该将 `self._insert()` 调用的结果分配给当前节点的 `left` 属性。

```js
({ test: () => assert.match(code, /node\.left\s*=\s*self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--    

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            pass
--fcc-editable-region--    

```
