---
id: 65c4f0d32ca17ad4b1636b0e
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

在 `__init__` 方法中，删除 `pass` 并将 `root` 初始化为值 `None`。 

`root` 属性表示二叉搜索树的根节点。 由于这是创建新的 `BinarySearchTree` 对象时的构造函数，它以一棵空树开始，因此 `root` 属性设置为 `None`。

# --hints--

你应该从 `__init__` 方法中删除 `pass` 语句。

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

你应该使用 `self.root` 将 `root` 属性初始化为 `None`。

```js
({ test: () => assert.match(code, /^\s{8}self\.root\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        pass
--fcc-editable-region--        
```
