---
id: 65c4f02dbd3557d23f12db62
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

在 `__init__` 方法内部，删除 `pass`，并使用 `self.key` 将 `key` 参数的值分配给节点的 `key` 属性。

这意味着 `TreeNode` 实例的 `key` 属性将被设置为对象创建期间传递的值。

# --hints--

你应该从 `__init__` 方法中删除 `pass` 语句。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("TreeNode").find_function("__init__").has_pass()`
      )
    );
  },
});

```

你应该使用 `self.key` 将 `key` 参数的值分配给节点的 `key` 属性。

```js
({ test: () => assert.match(code, /^\s{8}self\.key\s*=\s*key/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__(self, key):
        pass
--fcc-editable-region--
```
