---
id: 65c4f22498d22ed775ef8efb
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

现在需要检查 `node` 参数是否是 `None`。 如果是，则意味着该方法已到达应插入新节点的叶节点或树中的空点。

在 `_insert` 方法主体内，将 `pass` 替换为 `if` 语句，以检查是否 `node is None`。

在新的 `if` 块内，返回 `TreeNode(key)` 以使用提供的键创建一个新的 `TreeNode` 实例。 这将成为新的叶节点，有效地将键插入树中。

# --hints--

你的 `_insert` 方法中不应包含 `pass`。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_insert").has_pass()`)) })

```

你应该编写一个 `if` 语句来检查是否 `node is None`。

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^\s{8}if\s+node\s+is\s+None\s*:/m));
  }
})
```

你应该从 `if` 块返回 `TreeNode(key)`。

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^(\s{8})if\s+node\s+is\s+None\s*:\s*^\1\s{4}return\s+TreeNode\s*\(\s*key\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        pass
--fcc-editable-region--
```
