---
id: 65c4f09e074dd8d37830ea00
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

创建另一个名为 `BinarySearchTree` 的空类，代表二叉搜索树。

# --hints--

你应该有一个名为 `BinarySearchTree` 的类。 请记住使用 `pass` 关键字。

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(BinarySearchTree)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--
```
