---
id: 65c4f07da204f4d2f325fbd3
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

在 `__init__` 方法内部，将节点的 `left` 和 `right` 属性初始化为 `None`。 这是因为当第一次创建节点时，它没有任何左子节点或右子节点。 请记住使用 `self` 关键字。 

# --hints--

你应该使用 `self.left` 将 `None` 分配给节点的 `left` 属性。

```js
({ test: () => assert.match(code, /self\.left\s*=\s*None/m) })
```

你应该使用 `self.right` 将 `None` 分配给节点的 `right` 属性。

```js
({ test: () => assert.match(code, /self\.right\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:    
    def __init__(self, key):
        self.key = key
--fcc-editable-region--
```
