---
id: 65c4f013851cefd1a4fe4c96
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

`__init__` 方法接受两个参数：`self`（代表正在创建的类的实例）和 `key`（要存储在节点中的值）。 将这两个参数添加到 `__init__()` 方法。

# --hints--

你应该向该方法添加 `self` 参数。

```js
({ test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*/) });
```

你应该向该方法添加 `key` 参数。

```js
({
  test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__():
        pass
--fcc-editable-region--
```
