---
id: 65c4f2851645e8d84d14f5d1
title: 步骤11
challengeType: 20
dashedName: step-11
---

# --description--

现在你需要递归遍历树并使用二叉树的原理插入值：

- 小于键的值放在左子树中
- 大于键的值放在右子树中

在现有的条件语句之后，编写另一个 `if` 语句来检查 `key` 是否小于 `node.key`。

# --hints--

你应该编写另一个 if 语句来检查 `key` 是否小于 `node.key`。 记得使用 `pass` 关键字。

```js
({ test: () => assert(runPython(` _Node(_code).find_class("BinarySearchTree").find_function("_insert").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

--fcc-editable-region--
```

