---
id: 65c4eb814cc977c95cd7df0e
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在此项目中，你将创建二叉搜索树 (BST)。 BST 是一种数据结构，其中每个节点最多有两个子节点，左子节点包含小于父节点的值，右子节点包含大于父节点的值，从而可以进行高效的搜索和排序操作。

这是二叉搜索树的样子：

<img alt="a binary search tree" src="https://cdn.freecodecamp.org/curriculum/python/bst-example.png" style="background-color: white; height:500px; width:500px; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

首先定义一个空的 `TreeNode` 类。 `TreeNode` 类表示二叉搜索树中的节点。 使用 `pass` 关键字填充类主体来避免错误。

# --hints--

你应该使用 `class` 关键字声明一个名为 `TreeNode` 的空类。

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(TreeNode)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
