---
id: 65c4f33bc3c3d8db7f732787
title: 步骤 16
challengeType: 20
dashedName: step-16
---

# --description--

现在，为了执行实际的插入，在 `BinarySearchTree` 类中定义一个空的 `insert` 方法并为其提供 `self` 参数。

# --hints--

你应该在 `BinarySearchTree` 类中定义一个带有 `self` 参数的 `insert` 方法。 记住 `pass` 关键字。

```js
({ test: () => assert.match(code, /^\s{4}def\s+insert\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node
--fcc-editable-region--    

--fcc-editable-region--

```
