---
id: 65c4f3258d2e4cdacc919dfd
title: 步骤15
challengeType: 20
dashedName: step-15
---

# --description--

在 `_insert` 方法的末尾，插入过程完成后，返回当前节点以更新递归调用堆栈更高级别的树结构。


# --hints--

你应该返回条件块之外的当前节点。

```js
({ test: () => assert.match(code, /^\s{8}return\s+node/m) })
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--    
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)

--fcc-editable-region--

```

