---
id: 65c6409418806da8d0636ffc
title: 步骤 26
challengeType: 20
dashedName: step-26
---

# --description--

`insert` 和 `search` 功能已完成，现在是时候测试它们了。

请注意，此时，节点尚未排序，只是插入。 稍后你将使用中序遍历进行排序。

你可以像这样在 Python 中创建一个类的实例：

```python
object_name = ClassName()
```

在类定义之外，创建 `BinarySearchTree` 类的实例并将其分配给变量 `bst`。


# --hints--

你应该创建 `BinarySearchTree` 类的实例。

```js
assert.match(code, /BinarySearchTree\(\s*\)/);
```

你应该将 `BinarySearchTree` 的新实例分配给变量 `bst`。

```js
assert.match(code, /bst\s*=\s*BinarySearchTree\(\s*\)/);
``` 

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--


--fcc-editable-region--
```
