---
id: 65c646d4148ae3b2d1cbcac4
title: 第 30 步
challengeType: 20
dashedName: step-30
---

# --description--

请注意，你的查找返回类似 `80: <__main__.TreeNode object at 0x108b3e0>` 的内容。这是在打印一个类的实例时的默认字串表现。

要将其更改为打印有用的值，请在 `TreeNode` 类中定义另一个名为 `__str__` 的方法。 它接受一个参数 `self`。

定义 `__str__` 后，你会在控制台中收到异常，因为 `__str__` 方法尚未返回任何内容。 你将在下一步中处理方法主体。

# --hints--

你应该定义一个方法 `__str__`，它只接受一个参数 `self`。 记得使用 `pass`。

```js
assert.match(code, /^(\s+)def\s+__init__.+?^\1def\s+__str__\(\s*self\s*\)\s*:\s*\n^\1\1pass/ms)
```


# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

print('Search for 80:', bst.search(80))

```
