---
id: 65c63da2ed6769a10e141341
title: 步骤 20
challengeType: 20
dashedName: step-20
---

# --description--

现在你将为递归搜索定义一个基本情况。 删除当前的 `pass` 并编写一个检查两个条件的 `if` 语句：

- 如果 `node` 是 `None`：这表示搜索已到达分支的末尾，但尚未找到密钥。
- 如果是 `node.key == key`: 这意味着在当前节点中已经找到了该键。

将两个条件与 `or` 运算符结合起来，并返回 `if` 块内的当前节点。

# --hints--

你应该从 `_search` 方法中删除 `pass` 关键字。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_search").has_pass()`
      )
    );
  },
});

```

你应该编写一个 `if` 语句来检查 `node` 是否为 `None` 或者 `node.key` 是否等于 `key`。

```js
({ test: () => assert.match(code, /^\s{8}if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:/m) })
```

你应该从 `if` 块返回 `node`。

```js
({ test: () => assert.match(code, /^(\s{8})if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:\s*^\1\s{4}return\s+node/m) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)

--fcc-editable-region--
    def _search(self, node, key):
        pass

--fcc-editable-region--
```
