---
id: 65c644829cfb63acf3479d09
title: 第 28 步
challengeType: 20
dashedName: step-28
---

# --description--

通过遍历列表并对列表中的每个节点调用 `insert` 方法，将列表 `nodes` 中的节点插入到 `bst` 实例中。 使用 `node` 作为迭代变量。

# --hints--

你应该有一个 `for node in nodes` 循环。

```js
({ test: () => (runPython(`assert _Node(_code).find_for_loops()[0].find_for_vars().is_equivalent("node") and _Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("nodes")`)) });
```

你应该在 `for` 循环内有 `bst.insert(node)`。

```js
({ test: () => (runPython(`assert _Node(_code).find_for("node", "nodes").find_body().is_equivalent("bst.insert(node)")`)) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

--fcc-editable-region--
```
