---
id: 65c63e6962159fa372ecdbec
title: 步骤 22
challengeType: 20
dashedName: step-22
---

# --description--

如果第二个 `if` 语句不是 `True`，则表示目标键大于或等于当前节点键。 

在二叉搜索树中，如果目标键大于当前节点的键，则继续在右子树中搜索。

在 `if` 块之后，返回使用当前节点的右子节点和键作为参数调用 `_search` 方法的结果。

# --hints--

你应该使用 `self._search()` 递归调用 `_search` 方法。

```js
const after_split = code.split('return self._search(node.left, key)')[1];
assert.match(after_split, /self\._search/);
```

你应该向 `_search` 方法的调用传递两个参数：`node.right` 和 `key`。

```js
({ test: () => assert.match(code, /self\._search\(\s*node\.right\s*,\s*key\s*\)/) });
```

你应该返回 `_search` 方法调用的结果。

```js
({
  test: () => assert.match(code, /return\s+self\._search\(\s*node\.right\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)

--fcc-editable-region--
```
