---
id: 65c63df529bd15a24c187c62
title: 步骤21
challengeType: 20
dashedName: step-21
---

# --description--

编写另一个 `if` 语句，检查目标键是否小于当前节点的键。

在 `if` 块内，返回以当前节点的左子节点和 `key` 作为参数调用 `_search` 方法的结果。

# --hints--

你应该写另一个 `if` 语句来查看 `key` 是否小于 `node.key`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_search").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

你应该从新的 `if` 块返回 `self._search(node.left, key)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_search").find_ifs()[1].find_bodies()[0].is_equivalent("return self._search(node.left, key)")`)) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node

--fcc-editable-region--
```
