---
id: 65c63fd9b387f0a6c5bf6a72
title: 步骤 24
challengeType: 20
dashedName: step-24
---

# --description--

在 `search` 方法中，删除 `pass` 并使用以下参数调用辅助方法 `_search`。

- `self.root`：这是二叉搜索树的根。 搜索从根开始。
- `key`：这是用户想要在二叉搜索树中查找的值。

在内部，`search` 将实际的搜索逻辑委托给 `_search` 辅助方法，该方法在二叉搜索树中执行实际的递归搜索。

# --hints--

你应该从 `search` 方法中删除 `pass` 关键字。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("search").has_pass()`
      )
    );
  },
});

```

你应该在 `search` 方法中调用 `_search` 方法。

```js
({ test: () => assert(runPython(`
  str(_Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0]).startswith("self._search(")
`)) });
```

你应该调用 `_search` 方法，并将 `self.root` 和 `key` 作为参数传递。

```js
({ test: () => assert(runPython(`
  _Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0].is_equivalent("self._search(self.root, key)")
`)) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        pass
--fcc-editable-region--
```
