---
id: 65c645b838e7deb080fc25e0
title: 步骤 29
challengeType: 20
dashedName: step-29
---

# --description--

现在，节点已插入。 要查看它们是否已正确插入，你可以搜索树中的节点。

在 for 循环之外，在 `bst` 实例中搜索节点 `80` 并将其添加到 `print` 调用中。 此外，`print` 函数的第一个参数应该是 `'Search for 80:'`

# --hints--

你应该打印调用 `bst.search(80)` 的结果，并且你的 `print` 语句应该将第一个参数作为 `'Search for 80:'`。

```js
assert.match(code, /^print\(('|")Search for 80:('|"),\s*bst\.search\(80\)/gm);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

--fcc-editable-region--
```
