---
id: 65c63a4c4da62e9ae18e321a
title: 步骤 19
challengeType: 20
dashedName: step-19
---

# --description--

现在该研究搜索功能了。 就像为 `insert` 方法创建辅助方法 `_insert` 一样，你需要为 `search` 方法创建一个辅助方法 `_search`。

定义 `_search` 方法，其有三个参数，分别是 `self`、`node` 和 `key`。

# --hints--

你应该声明一个 `_search` 方法。 注意前导下划线并记得使用 `pass`。

```js
({ test: () => assert.match(code, /def\s+_search/) });
```

你的 `_search` 方法应该采用三个参数：`self`、`node` 和 `key`。

```js
({ test: () => assert.match(code, /def\s+_search\(\s*self\s*,\s*node\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
--fcc-editable-region--

--fcc-editable-region--
```
