---
id: 65ca03bd8eb5faf24b250c56
title: 步骤 33
challengeType: 20
dashedName: step-33
---

# --description--

下一步是进行节点的删除工作。 为此，你可以遵循相同的方法，首先定义辅助方法，然后定义实际方法。

在 `BinarySearchTree` 类中，定义一个名为 `_delete` 的新辅助方法，该方法接受三个参数：`self`、`node` 和 `key`。

# --hints--

你应该在 `BinarySearchTree` 类中定义 `_delete` 方法，并带有参数 `self`、`node` 和 `key`。 请记住使用 `pass` 关键字。

```js
({ test: () => assert.match(code, /def\s+_delete\(\s*self\s*,\s*node\s*,\s*key\s*\)/) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--


--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))

```
