---
id: 65ca075e3184180309f4c1d9
title: 步骤 39
challengeType: 20
dashedName: step-39
---

# --description--

在 `else` 块内，用 `if` 语句替换 `pass`，以检查当前节点的左子节点是否为 `None`。

当 `node.left` 是 `None` 时，没有左孩子。 因此，从新的 `if` 块中返回右子块作为替换。

# --hints--

你应该从 `else` 块中删除 `pass` 关键字。

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree").find_function("_delete"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});
```

你应该编写一个 `if` 语句来检查 `node.left` 是否是 `None`。

```js
assert.match(code, /if\s+node\.left\s*is\s*None/);
```

你应该返回当前节点的右子节点。

```js 
assert.match(code, /return\s+node\.right/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            pass   
--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
