---
id: 65cb45d130c97cb459439fac
title: 步骤 42
challengeType: 20
dashedName: step-42
---

# --description--

为了找到右子树中的最小值，你需要遍历给定节点的左子节点，直到到达子树中最左边（最小）的节点。

为此，编写一个 `while` 循环，该循环在 `node.left is not None` 时运行，并将 `pass` 移动到 `while` 块内。 此条件检查是否存在左孩子。 只要有左孩子，循环就继续，并会找到一个较小的值。

# --hints--

你应该从 `_min_value` 方法中删除 `pass` 关键字。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_min_value").has_pass()`
      )
    );
  },
});
```

你应该在 `while` 循环中使用条件 `node.left is not None`。

```js
({ test: () => assert.match(code, /while\s+node\.left\s+is\s+not\s+None\s*:/) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left  

--fcc-editable-region--
    def _min_value(self, node):
        pass

--fcc-editable-region--   

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
