---
id: 65d8a6fcb15a3a239ba35dfd
title: 步骤46
challengeType: 20
dashedName: step-46
---

# --description--

找到最小值后，你需要递归地从右子树中删除具有最小值的节点。

这一步确保了从树中删除具有最小值的节点，同时保持二叉搜索树（BST）属性。

使用 `node.right` 和 `node.key` 作为参数递归调用 `_delete` 方法。 将 `_delete()` 调用的返回值分配给当前节点的右子节点。

# --hints--

你应该用 `node.right` 和 `node.key` 作为参数递归地调用 `_delete` 方法。

```js
assert.match(code, /self\._delete\(\s*node\.right\s*,\s*node\.key\s*/);
```

你应该将 `_delete()` 调用的返回值分配给当前节点的右子节点。

```js
assert.match(code, /node\.right\s*=\s*self\._delete\(\s*node\.right\s*,\s*node\.key/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key)
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left

            node.key = self._min_value(node.right) 
            
--fcc-editable-region--
    def _min_value(self, node):
        while node.left is not None:
            node = node.left

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
