---
id: 66ad984e92ce213a66f7e51c
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Тепер настав час призначити значення вашій змінній `bot`.

У попередніх уроках ви навчились призначати значення змінним ось так:

```js
variableName = "Here is the value";
```

Пам’ятайте, що те, що знаходиться праворуч від оператора присвоєння `=`, — це значення, яке ви призначаєте змінній зліва.

Призначте рядок `"teacherBot"` змінній `bot`.

# --hints--

Ви повинні призначити рядок своїй змінній `bot`.

```js
assert.isString(bot);
```

Вашій змінній `bot` має бути призначений рядок `"teacherBot"`.

```js
assert.equal(bot, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
