---
id: 66ad998285df023b23dacdd3
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Тепер настав час ініціалізувати змінну `botLocation`.

Коли потрібно оголосити змінні з кількома словами, можна використовувати конвенцію найменування <dfn>camelCase</dfn>.

При використанні `camelCase` перше слово пишеться повністю в нижньому регістрі, а перша літера кожного наступного слова — з великої.

Ось приклад:

```js
let thisIsCamelCase = "Alice the camel is actually a horse.";
```

Оголосіть і призначте рядок `"the universe"` змінній `botLocation` в одному рядку, використовуючи ключове слово `let`.

# --hints--

Використовуйте ключове слово `let` для оголошення змінної.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

Ваша змінна має називатися `botLocation`.

```js
assert.isNotNull(botLocation);
```

Ви повинні призначити рядок змінній `botLocation`.

```js
assert.isString(botLocation);
```

Змінній `botLocation` має бути призначений рядок `"the universe"`.

```js
assert.equal(botLocation, "the universe");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
