---
id: 66ada3f46945763dd97f43f8
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Раніше ви створили змінні `bot` та `botLocation`. Тепер ви використаєте їх, щоб вивести нові повідомлення в консоль.

У попередніх уроках ви навчились працювати з конкатенацією рядків, використовуючи оператор `+` для об’єднання рядків ось так:

```js
let firstName = "John";

console.log("Hello, my name is " + firstName + ".");
// result: "Hello, my name is John."
```

Пам’ятайте, що потрібно уважно ставитися до пробілів при конкатенації рядків зі змінними.

Створіть змінну з назвою `botIntroduction`.

Потім використайте конкатенацію рядків з оператором `+`, щоб об’єднати рядок `"My name is "`, за яким іде змінна `bot`, а потім крапка (`.`).

Призначте це значення змінній `botIntroduction`.

Потім виведіть у консоль змінну `botIntroduction`.

# --hints--

У вас має бути змінна з назвою `botIntroduction`.

```js
assert.isNotNull(botIntroduction);
```

Значення вашої змінної `botIntroduction` має бути рядком.

```js
assert.isString(botIntroduction);
```

Використайте конкатенацію рядків з оператором `+`, щоб об’єднати рядок `"My name is "` зі змінною `bot`, а потім крапку (`.`). Будьте уважні до пробілів. Призначте цей об’єднаний рядок змінній `botIntroduction`.

```js
assert.equal(botIntroduction, "My name is teacherBot.");
```

У вашому коді має бути четверта інструкція `console.log()`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 4);
```

Інструкція `console` має виводити змінну `botIntroduction`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botIntroduction\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

--fcc-editable-region--

--fcc-editable-region--
```
