---
id: 66adc82170acb1464bd348eb
title: Крок 13
challengeType: 1
dashedName: step-13
---

# --description--

Останні кілька повідомлень від бота будуть зосереджені на улюбленій темі бота.

Почніть зі створення нової змінної з назвою `favoriteSubject` і призначте їй рядок `"Computer Science"`.

# --hints--

У вас має бути змінна з назвою `favoriteSubject`.

```js
assert.isNotNull(favoriteSubject);
```

Ваша змінна `favoriteSubject` має бути рядком.

```js
assert.isString(favoriteSubject);
```

Ви повинні призначити рядок `"Computer Science"` змінній `favoriteSubject`.

```js
assert.strictEqual(favoriteSubject, "Computer Science");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

--fcc-editable-region--

--fcc-editable-region--
```
