---
id: 66ad8d150264db3926eccfeb
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

У попередніх уроках ви дізналися про ключове слово `let` і як оголосити змінні в JavaScript.

Ось нагадування, як оголосити змінну за допомогою ключового слова `let`:

```js
let greeting;
```

Використайте ключове слово `let`, щоб оголосити змінну з назвою `bot`.

**ПРИМІТКА**: Ви використовуєте `let` тут, тому що пізніше на практичному занятті ви будете змінювати значення змінної `bot`.

# --hints--

У вашому коді має бути ключове слово `let`.

```js
assert.match(__helpers.removeJSComments(code), /let/g);
```

Ви маєте оголосити змінну з назвою `bot`.

```js
assert.isNotNull(bot);
```

Ви маєте використати ключове слово `let`, щоб оголосити змінну.

```js
assert.match(__helpers.removeJSComments(code), /let\s+bot;?/g);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
--fcc-editable-region--

--fcc-editable-region--
```
