---
id: 66ad8ab945d266383f318cbf
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Тепер ви повинні побачити перше повідомлення від бота в консолі.

Настав час додати друге повідомлення від бота.

Додайте ще одну інструкцію `console.log`, щоб вивести повідомлення `"I am excited to talk to you."` у консоль.

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

У вашому коді має бути друга інструкція `console.log()`.

```js
assert.lengthOf(logSpy.calls, 2);
```

Ваша інструкція `console` має містити повідомлення `"I am excited to talk to you."`. Не забудьте лапки.

```js
assert.isTrue(
  logSpy.calls.some(
    (call) => /I\s+am\s+excited\s+to\s+talk\s+to\s+you\s*\./.test(call[0])
  )
);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
--fcc-editable-region--

--fcc-editable-region--
```
