---
id: 66adcf383276814776aba3ca
title: Крок 14
challengeType: 1
dashedName: step-14
---

# --description--

Далі створіть змінну з назвою `favoriteSubjectSentence`.

Використайте конкатенацію рядків, щоб об’єднати рядок `"My favorite subject is "` зі змінною `favoriteSubject`, після чого додайте крапку.

Призначте результат змінній `favoriteSubjectSentence`.

Потім виведіть значення `favoriteSubjectSentence` у консоль.

# --hints--

У вас має бути змінна з назвою `favoriteSubjectSentence`.

```js
assert.isNotNull(favoriteSubjectSentence);
```

Змінна `favoriteSubjectSentence` має бути рядком.

```js
assert.isString(favoriteSubjectSentence);
```

Використайте конкатенацію рядків, щоб об’єднати рядок `"My favorite subject is "` зі змінною `favoriteSubject`, після чого додайте крапку. Зверніть увагу на пробіли.

```js
assert.strictEqual(favoriteSubjectSentence, "My favorite subject is Computer Science.");
```

У вашому коді має бути інструкція `console`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 8);
```

Виведіть значення `favoriteSubjectSentence` у консоль.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*favoriteSubjectSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

const favoriteSubject = "Computer Science";

--fcc-editable-region--

--fcc-editable-region--
```
