---
id: 66ad8294a0ad902f1b31b612
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви навчитеся працювати з основами JavaScript, створюючи вітального бота.

На цьому першому кроці вам потрібно вивести повідомлення в консоль від імені вітального бота.

Пам’ятайте, що ви вже вивчали `console.log()` і рядки в попередніх уроках.

Ось нагадування, як використовувати `console.log()` з рядками:

```js
console.log("Hello, World!");
```

Додайте інструкцію `console.log()`, яка виводить рядок `"Hi there!"` у консоль. Не забудьте лапки навколо повідомлення!

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

У вашому коді має бути інструкція `console.log()`.

```js
assert.isNotEmpty(logSpy.calls);
```

Ваша інструкція `console` має містити повідомлення `"Hi there!"`. Не забудьте лапки.

```js
assert.deepInclude(logSpy.calls, ["Hi there!"]);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
