---
id: 66adb844118ba74107ce771f
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Наступне повідомлення від бота стосуватиметься місцезнаходження бота.

Створіть змінну з назвою `botLocationSentence`.

Потім використайте конкатенацію рядків з оператором `+`, щоб об’єднати рядок `"I live in "` зі змінною `botLocation`, після чого додайте крапку (`.`).

Призначте це значення змінній `botLocationSentence`.

Потім виведіть у консоль значення `botLocationSentence`.

# --hints--

У вас має бути змінна `botLocationSentence`.

```js
assert.isNotNull(botLocationSentence);
```

Ваша змінна `botLocationSentence` має містити значення рядка.

```js
assert.isString(botLocationSentence);
```

Ви маєте використати конкатенацію рядків з оператором `+`, щоб об’єднати рядок `"I live in "` зі змінною `botLocation`, після чого додати крапку (`.`). Зверніть увагу на пробіли.

```js
assert.equal(botLocationSentence, "I live in the universe.");
```

У вашому коді має бути інструкція `console`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 5);
```

Ваша інструкція `console` має виводити змінну `botLocationSentence`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botLocationSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
