---
id: 66adc42868cab843ccee87d9
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

Тепер настав час побачити нове значення `bot`.

Почніть зі створення нової змінної з назвою `nicknameIntroduction`.

Використайте конкатенацію рядків, щоб об’єднати рядок `"My nickname is "` зі змінною `bot`, після чого додайте крапку (`.`).

Призначте отриманий рядок змінній `nicknameIntroduction`.

Потім виведіть значення `nicknameIntroduction` у консоль.

# --hints--

У вас має бути змінна з назвою `nicknameIntroduction`.

```js
assert.isNotNull(nicknameIntroduction);
```

Ваша змінна `nicknameIntroduction` має бути рядком.

```js
assert.isString(nicknameIntroduction);
```

Ви маєте використати конкатенацію рядків, щоб об’єднати рядок `"My nickname is "` зі змінною `bot`, після чого додати крапку (`.`). Зверніть увагу на пробіли.

```js
assert.strictEqual(nicknameIntroduction, "My nickname is professorBot.");
```

У вашому коді має бути інструкція `console`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 6);
```

Ви маєте вивести значення `nicknameIntroduction` у консоль.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*nicknameIntroduction\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

--fcc-editable-region--

--fcc-editable-region--
```
