---
id: 66ad9c633fa26d3c1475eae3
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Тепер настав час додати ще одне повідомлення від бота.

Додайте ще одну інструкцію `console` у код, яка виводить повідомлення `"Allow me to introduce myself."`.

# --hints--

У вашому коді має бути третя інструкція `console.log()`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 3);
```

Ваша інструкція `console` має містити повідомлення `"Allow me to introduce myself."`. Не забудьте лапки.

```js
assert.match(__helpers.removeJSComments(code), /("|')Allow\s+me\s+to\s+introduce\s+myself.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
