---
id: 66ad9c633fa26d3c1475eae3
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Тепер настав час додати ще одне повідомлення від бота.

Додайте ще одну інструкцію `console` у код, яка виводить повідомлення `"Allow me to introduce myself."`.

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

У вашому коді має бути третя інструкція `console.log()`.

```js
assert.lengthOf(logSpy.calls, 3);
```

Ваша інструкція `console` має містити повідомлення `"Allow me to introduce myself."`. Не забудьте лапки.

```js
assert.isTrue(
  logSpy.calls.some(
    (call) => /Allow\s+me\s+to\s+introduce\s+myself\s*\./.test(call[0])
  )
);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
