---
id: 66adc6046acc18453decf577
title: Крок 12
challengeType: 1
dashedName: step-12
---

# --description--

Щоб побачити новий нікнейм бота, потрібно вивести нове повідомлення в консоль.

Створіть нову змінну з назвою `newNicknameGreeting`.

Потім використайте конкатенацію рядків, щоб об’єднати рядок `"I love my nickname but I wish people would call me "` зі змінною `bot`, додавши в кінці крапку.

Призначте результат змінній `newNicknameGreeting`.

Потім виведіть значення `newNicknameGreeting` у консоль.

# --hints--

У вас має бути змінна з назвою `newNicknameGreeting`.

```js
assert.isNotNull(newNicknameGreeting);
```

Значення змінної `newNicknameGreeting` має бути рядком.

```js
assert.isString(newNicknameGreeting);
```

Використайте конкатенацію рядків, щоб об’єднати рядок `"I love my nickname but I wish people would call me "` зі змінною `bot`, додавши в кінці крапку.

```js
assert.strictEqual(newNicknameGreeting, "I love my nickname but I wish people would call me awesomeTeacherBot.");
```

У вашому коді має бути інструкція `console`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 7);
```

Виведіть значення `newNicknameGreeting` у консоль.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*newNicknameGreeting\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
