---
id: 66629f407d679d3105e8317f
title: 建置一個剪刀石頭布遊戲
challengeType: 25
dashedName: top-build-a-rock-paper-scissors-game
---

# --description--
你的遊戲將與電腦對戰。你將撰寫一個函式（程式），隨機傳回 `"rock"`、`"paper"` 或 `"scissors"`。

你不需要擔心遊戲的前端部分。你只需要撰寫遊戲的邏輯。打開 `script.js` 並開始按照指示操作。

**使用者故事：**

1. 你應該有一個名為 `getComputerChoice` 的函式（程式）。

1. 你的 `getComputerChoice` 函式（程式）應該隨機傳回 `"rock"`、`"paper"` 或 `"scissors"`。

**提示：** [Math.random](https://developer.mozilla.org/en-US/docs/Web/JavaScript/Reference/Global_Objects/Math/random) 方法會傳回一個大於或相等於 0 且小於 1 的隨機數。想想你如何利用這點來有條件地傳回多個選項中的其中一個。

你的遊戲將由真人玩家進行。你將撰寫一個函式，該函式接收使用者的選擇並傳回它。

1. 創建一個名為 `getHumanChoice` 的函式（程式）。

1. 撰寫程式碼，使 `getHumanChoice` 根據使用者輸入傳回其中一個有效選項。

**提示：** 使用 [prompt](https://developer.mozilla.org/en-US/docs/Web/API/Window/prompt) 方法來取得使用者的輸入。

你的遊戲會追蹤玩家的分數。你會撰寫變數來追蹤玩家的分數。

1. 在全域的範圍中創建兩個名為 `humanScore` 和 `computerScore` 的新變數。

1. 將那些變數初始化為 `0`。

你的遊戲將會一輪一輪地進行。你將撰寫一個函式（程式），該函式以人類和電腦玩家的選擇作為引數，執行單輪遊戲，增加該輪勝利者的分數，並記錄勝利者公告。

1. 創建一個名為 `playRound` 的新函式（程式）。

1. 為 `playRound` 定義兩個參數。第一個參數為 `humanChoice`，第二個參數為 `computerChoice`。使用這兩個參數將人類和電腦的選擇作為引數。

1. 讓你的函式（程式）`humanChoice` 參數不區分大小寫，這樣玩家就可以輸入 `"rock"`、`"ROCK"`、`"RocK"` 或其他變體。

1. 撰寫你的 `playRound` 函式（程式），該函式（程式）傳回一個字串值，表述該回合的勝利者。

- 如果是平手，應該傳回 `"It's a tie!"`。

- 如果玩家贏了，應該傳回 `"You win! [player choice] beats [computer choice]"`。

- 如果電腦贏了，應該傳回 `"You lose! [computer choice] beats [player choice]"`。

1. 根據回合勝者遞增 `humanScore` 或 `computerScore` 變數。

你的遊戲將進行 3 輪。你將撰寫一個名為 `playGame` 的函式（程式），該函式會呼叫 `playRound` 來進行 3 輪，追蹤分數，並在結束時宣告勝利者。

1. 創建一個名為 `playGame` 的新函式（程式）。

1. 創建一個迴圈，執行 3 輪，並每次以人的選擇和電腦的選擇函式作為引數呼叫 `playRound` 函式（程式）。

1. 在遊戲結束時，根據誰贏得最多回合來傳回遊戲的贏家。

- 如果真人玩家贏得的回合數比電腦玩家多，傳回一則訊息，內容為 `"You win the game!"`。

- 如果電腦玩家贏得的回合數比人類玩家多，傳回一則訊息，內容為 `"You lose the game!"`。

# --hints--

你應該有一個名為 `getComputerChoice` 的函式（程式）。

```js
assert.isFunction(getComputerChoice);
```

你的 `getComputerChoice` 函式（程式）應該隨機傳回 `"rock"`、`"paper"` 或 `"scissors"`。

```js
const counts = {}
for (let i = 0; i < 100; i++) {
  const result = getComputerChoice();
  counts[result] = (counts[result] ?? 0) + 1;
}
assert.lengthOf(Object.keys(counts), 3);
assert.include(Object.keys(counts), "rock");
assert.include(Object.keys(counts), "paper"); 
assert.include(Object.keys(counts), "scissors");
```

你應該有一個名為 `getHumanChoice` 的函式（程式）。

```js
assert.isFunction(getHumanChoice);
```

你應該在全域的範圍內有兩個名為 `humanScore` 和 `computerScore` 的變數。

```js
assert.exists(humanScore);
assert.exists(computerScore);
```

你應該有一個名為 `playRound` 的函式（程式）。

```js
assert.isFunction(playRound);
```

你的 `playRound` 函式（程式）應該以 `humanChoice` 和 `computerChoice` 參數接收人類和電腦玩家的選擇作為引數。

```js
assert.match(playRound.toString(), /\s*(?:\bhumanChoice\b\s*,\s*\bcomputerChoice\b)/);
```

你的 `playRound` 函式（程式）應該不區分大小寫，讓玩家可以輸入 `"rock"`、`"ROCK"`、`"RocK"` 或其他變體。

```js
assert.match(playRound.toString(), /\bhumanChoice\s*\.toLowerCase\(\)/);
```

你的 `playRound` 函式（程式）應該在使用者和電腦選擇相同時傳回字串 `"It's a tie!"`。

```js
assert.equal(playRound("rock", "rock"), "It's a tie!");
```

你的 `playRound` 函式（程式）應該傳回字串 `"You win! [player choice] beats [computer choice]"`，如果玩家獲勝。

```js
assert.equal(playRound("rock", "scissors"), "You win! rock beats scissors");
```

你的 `playRound` 函式（程式）應該傳回字串 `"You lose! [computer choice] beats [player choice]"`，如果電腦贏了。

```js
assert.equal(playRound("rock", "paper"), "You lose! paper beats rock");
```

你的 `playRound` 函式應該根據回合贏家遞增 `humanScore` 或 `computerScore` 變數。

```js
humanScore = 0;
computerScore = 0;
playRound("rock", "scissors");
assert.equal(humanScore, 1);
```

你應該有一個名為 `playGame` 的函式（程式）。

```js
assert.isFunction(playGame);
```

你應該使用迴圈來執行 3 輪。

```js
assert.match(playGame.toString(), /\bfor\s*\(/);
```

你應該根據誰贏得最多回合來傳回遊戲的贏家。

```js
window.prompt = () => "rock";
assert.match(playGame(), /You (win|lose) the game!/);
```

# --seed--

## --seed-contents--

```js
const hand = ['rock', 'paper', 'scissors'];
```

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

# --solutions--

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

```js
const hand = ['rock', 'paper', 'scissors'];


let computerScore = 0;
let humanScore = 0;

const getComputerChoice = () => {
  return hand[Math.floor(Math.random() * hand.length)];
}

const getHumanChoice = () => {
  return prompt();
}

const playRound = (humanChoice, computerChoice) => {
  humanChoice = humanChoice.toLowerCase()

  const tie = "It's a tie!"
  const win = `You win! ${humanChoice} beats ${computerChoice}`
  const lose = `You lose! ${computerChoice} beats ${humanChoice}`


  if (humanChoice == 'rock') {
    if (computerChoice == 'rock') {
      computerScore++;
      humanScore++;
      return tie;
    } else if (computerChoice == 'paper') {
      computerScore++;
      return lose;
    } else {
      humanScore++;
      return win;
    }
  } else if (humanChoice == 'paper') {
    if (computerChoice == 'rock') {
      humanScore++;
      return win;
    } else if (computerChoice == 'paper') {
      computerScore++;
      humanScore++;
      return tie;
    } else {
      computerScore++;
      return lose;
    }
  } else if (humanChoice == 'scissors') {
    if (computerChoice == 'rock') {
      computerScore++;
      return lose;
    } else if (computerChoice == 'paper') {
      humanScore++;
      return win;
    } else {
      computerScore++;
      humanScore++;
      return tie;
    }
  }
}

const playGame = () => {
  for(let i = 0; i < 3; i++){
    playRound(getHumanChoice(), getComputerChoice())
  }

  return humanScore > computerScore ? "You win the game!" : "You lose the game!"
}

```

