---
id: 6995d0325fd7da40854fac40
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In una lezione precedente, hai imparato la ricerca binaria e come si differenzia dalla ricerca lineare.

In questo workshop, metterai in pratica questo concetto costruendo un algoritmo di ricerca binaria.

Per cominciare, crea una funzione `binarySearch` con due parametri `searchList` e `value`.

`searchList` è l'array di elementi all'interno del quale la funzione cercherà, e `value` è l'obiettivo che la funzione proverà a trovare all'interno di `searchList`.

# --hints--

Dovresti creare una funzione `binarySearch`.

```js
assert.isFunction(binarySearch);
```

La tua funzione `binarySearch` dovrebbe avere due parametri `searchList` e `value`.

```js
const binSearchParams = __helpers.getFunctionParams(binarySearch.toString());

const binarySearchParamsNames = binSearchParams.map(param => param.name)
assert.include(binarySearchParamsNames, "searchList");
assert.include(binarySearchParamsNames, "value");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
