---
id: 6996cfb6c01e9a802add38be
title: Passaggio 13
challengeType: 1
dashedName: step-13
---

# --description--

Infine, se il `value` non si trova al centro e non è nemmeno maggiore di `valueAtMiddle`, allora il `value` deve essere minore di `valueAtMiddle` e deve trovarsi a sinistra. Il blocco finale `else` gestirà questo caso.

Quindi, completa il loop rimuovendo la parola chiave `break` e aggiungendo un blocco `else`. Al suo interno, aggiorna la variabile `high` sottraendo `1` da `mid`.

# --hints--

Dovresti sostituire la parola chiave `break` con un blocco `else` che aggiorna la variabile `high` a `mid - 1`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s*\{\s*high\s*=\s*mid\s*-\s*1\s*;?\s*\}/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    }

--fcc-editable-region--      
    break;
--fcc-editable-region--

  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
