---
id: 6996d18b31aab5cbbffbc731
title: Passaggio 16
challengeType: 1
dashedName: step-16
---

# --description--

Un ultimo piccolo miglioramento che puoi aggiungere è mostrare anche l’indice in cui si trova l’obiettivo `value`.

Quindi, invece di restituire solo `pathToTarget` nell'`if`, restituisci un array contenente `pathToTarget` e un letterale template con il messaggio `Value found at index ${mid}`.

Con questo, il tuo workshop sull’algoritmo di ricerca binaria è completo!

# --hints--

Dovresti aggiornare l'istruzione `return` nell'`if` in modo che la funzione restituisca un array con `pathToTarget` come primo elemento e un letterale template per `Value found at index ${mid}` come secondo elemento. Assicurati di usare i backtick per il letterale template.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch: binarySearchFn } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearchFn?.toString());
assert.match(
  codeWithoutComments,
  /return\s*\[\s*pathToTarget\s*,\s*(?:`Value found at index\s*\$\{\s*mid\s*\}`|['"]Value found at index\s*['"]\s*\.\s*concat\(\s*mid\s*\))\s*\]/
);
```

La tua funzione dovrebbe restituire i risultati attesi.

```js
assert.deepEqual(
  binarySearch([1, 2, 3, 4, 5], 3), [[3], "Value found at index 2"]
);
assert.deepEqual(
  binarySearch([1, 2, 3, 4, 5, 9], 4), [[3, 5, 4], "Value found at index 3"]
);
assert.deepEqual(
  binarySearch([1, 3, 5, 9, 14, 22], 10), [[], "Value not found"]
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
--fcc-editable-region--
      return pathToTarget;
--fcc-editable-region--
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [[], "Value not found"];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```

# --solutions--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return [pathToTarget, `Value found at index ${mid}`];
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [[], "Value not found"];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```
