---
id: 6996c675b599baa7b103192d
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Nella ricerca binaria, esiste l’indice più basso possibile e l’indice più alto possibile.

L’indice più basso possibile rappresenta il limite più a sinistra nella porzione corrente dei valori in cui si sta cercando, in questo caso `searchList`. L’indice più alto possibile, invece, rappresenta il limite più a destra nei valori in cui si sta cercando.

Per tenere conto di questi due, devi considerare l’intero array in cui si sta cercando. Quindi, all’interno della funzione `binarySearch`, definisci una variabile `low` con valore `0` e una variabile `high` con un valore che tenga conto dell’ultimo indice nell’array in cui si sta cercando.

# --hints--

Dovresti avere una variabile chiamata `low` impostata su `0` nella tua funzione `binarySearch`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /(const|let|var)\s+low\s*=\s*0/);
```

Dovresti avere una variabile chiamata `high` impostata su `searchList.length - 1` nella tua funzione `binarySearch`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+high\s*=\s*searchList\.length\s*-\s*1/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
