---
id: 6996d0a2c2ac55041145a13f
title: Passaggio 14
challengeType: 1
dashedName: step-14
---

# --description--

Hai notato che la seconda chiamata della funzione ora ha i valori `3, 5, 4` nella lista? Questo indica che la ricerca binaria sta funzionando come previsto.

Ecco come è successo: l'algoritmo ha prima controllato `3` come centro della lista iniziale. Poiché `4` è maggiore di `3`, la ricerca si è spostata nella metà destra. Poi ha esaminato `5` come nuovo centro. Poiché `4` è minore di `5`, la ricerca si è spostata a sinistra, identificando infine `4` come centro dell'intervallo finale.

Per testare di nuovo la funzione, chiamala con `[1, 3, 5, 9, 14, 22], 10` e stampala nella console. Questa è una situazione in cui il valore non verrà trovato.

# --hints--

Dovresti chiamare la funzione con `binarySearch([1, 3, 5, 9, 14, 22], 10)` e stamparla nella console.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*3\s*,\s*5\s*,\s*9\s*,\s*14\s*,\s*22\s*\]\s*,\s*10\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
--fcc-editable-region--

--fcc-editable-region--
```
