---
id: 6996c60281feca6718cc15c1
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

All'interno della tua `binarySearch` funzione, crea una variabile `pathToTarget` e inizializzala con un array vuoto, per cominciare.

# --hints--

Dovresti avere una variabile `pathToTarget` impostata su un array vuoto `[]` all'interno della tua funzione `binarySearch`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+pathToTarget\s*=\s*\[\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
