---
id: 6996ccadf9048c79b857bb15
title: Passaggio 11
challengeType: 1
dashedName: step-11
---

# --description--

Puoi vedere che la prima chiamata restituisce `[3]`. Questo perché `3` è il valore centrale in `[1, 2, 3, 4, 5]`.

D'altra parte, la seconda chiamata restituisce un array vuoto. Questo succede perché, dopo aver controllato il valore centrale, il loop attualmente non ha la logica per restringere l'intervallo di ricerca aggiornando le variabili `low` o `high`. Funziona solo come un controllo singolo del punto medio e poi termina senza ulteriori ricerche.

Per permettere alla ricerca binaria di continuare a restringere la ricerca, aggiungi un'istruzione `else if` che controlli se `value` è maggiore di `valueAtMiddle`.

# --hints--

Dovresti aggiungere un'istruzione `else if` con la condizione `value > valueAtMiddle`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s+if\s*\(\s*value\s*>\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    if (value === valueAtMiddle) {
      return pathToTarget;
    } 
--fcc-editable-region--
    break;
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
