---
id: 6996c60281feca6718cc15c1
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

All'interno della tua `binarySearch` funzione, crea una variabile `pathToTarget` e inizializzala con un array vuoto, per cominciare.

# --hints--

Dovresti dichiarare una variabile `pathToTarget` impostata su un array vuoto all'interno della tua funzione `binarySearch`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { pathToTarget } = binarySearch?.variables ?? {};
const value = eval(pathToTarget?.value.toString())
assert.isArray(value);
assert.isEmpty(value);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
