---
id: 6996c675b599baa7b103192d
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Nella ricerca binaria, esiste l'indice più basso possibile e l'indice più alto possibile.

L’indice più basso possibile rappresenta il limite più a sinistra nella porzione corrente dei valori in cui si sta cercando, in questo caso `searchList`. L’indice più alto possibile, invece, rappresenta il limite più a destra nei valori in cui si sta cercando.

Per tenere conto di questi due, devi considerare l'intero array che stai cercando. Quindi, all'interno della funzione `binarySearch`, dichiara una variabile `low` con valore `0` e una variabile `high` con un valore che rappresenta l'ultimo indice dell'array in cui stai cercando.

# --hints--

Dovresti dichiarare una variabile chiamata `low` impostata a `0` nella tua funzione `binarySearch`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { low } = binarySearch?.variables ?? {};
assert.isTrue(low?.value.matches("0"));
```

Dovresti dichiarare una variabile chiamata `high` impostata a `searchList.length - 1` nella tua funzione `binarySearch`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { high } = binarySearch?.variables ?? {};
assert.isTrue(high?.value.matches("searchList.length - 1"));
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
