---
id: 6996c796385b6ea4e222d5e8
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Per iniziare con il loop, devi trovare l'indice centrale dello spazio di ricerca attuale.

Puoi ottenere l'indice centrale calcolando la media delle variabili `low` e `high`. Dovrai usare la divisione intera per ottenere la media dopo aver sommato i due valori, così la risposta sarà sempre un numero intero arrotondato per difetto, dato che gli indici devono essere numeri interi.

All'interno del tuo loop `while`, crea una variabile `mid` impostata sulla media di `low` e `high`.

# --hints--

All'interno del while loop dovresti creare una variabile `mid` impostata su `Math.floor((low + high) / 2)`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+mid\s*=\s*Math\.floor\(\s*\(\s*low\s*\+\s*high\s*\)\s*\/\s*2\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
--fcc-editable-region--
      
--fcc-editable-region--
  }
}
```
