---
id: 6996c7363a4a437839f1f15d
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Ora che hai definito i limiti della tua ricerca, è il momento di creare il loop che eseguirà l'algoritmo di ricerca binaria.

La ricerca binaria funziona restringendo ripetutamente lo spazio di ricerca. Questo processo continua finché esiste un intervallo valido di elementi da controllare.

Puoi esprimere questo con un loop `while` che continuerà finché il puntatore `low` è minore o uguale al puntatore `high`.

Se `low` diventa mai maggiore di `high`, significa che lo spazio di ricerca è diventato vuoto e `value` non si trova nell'array.

All'interno della funzione `binarySearch`, crea un loop `while` con una condizione che verifica se `low` è minore o uguale a high.

# --hints--

Dovresti creare un loop `while` che gira mentre `low` è minore o uguale a `high`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /while\s*\(\s*low\s*<=\s*high\s*\)/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
