---
id: 6996c8375f4301004459ee85
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Ora che hai calcolato l'indice `mid`, devi effettivamente recuperare il valore dal `searchList` a quell'indice. Questo valore è quello che confronterai con il parametro `value`, che è l'obiettivo della tua ricerca.

Crea una variabile `valueAtMiddle` e assegnale l'elemento da `searchList` che si trova all'indice `mid`.

# --hints--

Dovresti creare una variabile `valueAtMiddle` impostata su `searchList[mid]`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearch?.toString());
assert.match(
  codeWithoutComments,
  /(const|let|var)\s+valueAtMiddle\s*=\s*searchList\s*\[\s*mid\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
