---
id: 6996c95e5bede03bdf2937a2
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Il passaggio successivo è creare una condizione che verifichi se l'obiettivo `value` si trova al centro.

All'interno del loop `while`, crea un'istruzione `if` che controlli se l'obiettivo `value` è uguale a `valueAtMiddle`. Se lo è, restituisci la variabile `pathToTarget`.

# --hints--

Dovresti creare un'istruzione `if` con la condizione `value === valueAtMiddle`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)/
);
```

Dovresti restituire la variabile `pathToTarget` all'interno dell'istruzione `if`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?return\s+pathToTarget/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
