---
id: 6996cb5869acf58de751ae81
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

Ricorda che la funzione prende un array chiamato `searchList` e il `value` che stai cercando.

Ora, chiama la funzione con `binarySearch([1, 2, 3, 4, 5], 3)` e `binarySearch([1, 2, 3, 4, 5, 9], 4)` e registra entrambe le istruzioni nella console.

# --hints--

Dovresti chiamare la funzione con `binarySearch([1, 2, 3, 4, 5], 3)` e registrarla nella console.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*2\s*,\s*3\s*,\s*4\s*,\s*5\s*\]\s*,\s*3\s*\)\s*\)/
);
```

Dovresti chiamare la funzione con `binarySearch([1, 2, 3, 4, 5, 9], 4)` e registrarla nella console.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*2\s*,\s*3\s*,\s*4\s*,\s*5\s*,\s*9\s*\]\s*,\s*4\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
    break;
  }
  return [];
}

--fcc-editable-region--  

--fcc-editable-region--  
```
