---
id: 6996ca21a74104721f2c7e5d
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

Devi testare alcune cose per capire il flusso dell’algoritmo in questa fase iniziale.

Per farlo, devi prima uscire dal loop. Questo perché l’implementazione attuale permette solo un’iterazione, quindi se la condizione non è soddisfatta, si creerà un loop infinito.

Subito dopo l’istruzione `if`, usa la parola chiave `break` per uscire dal loop `while`. Poi, dopo il loop `while`, restituisci un array vuoto per indicare che il valore non è stato trovato.

# --hints--

Dovresti usare la parola chiave `break` per uscire dal loop `while` dopo l’istruzione `if`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?\}\s*break\s*;?/
);
```

Dovresti restituire un array vuoto dopo il tuo loop `while`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /return\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--
    
  }
  
--fcc-editable-region--
}
```
