---
id: 6996c8f94fe6b250f11a2a1e
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Ora, includi `valueAtMiddle` nell'array `pathToTarget` per tracciare i passaggi effettuati durante la ricerca, indipendentemente dal fatto che sia l'obiettivo. Puoi farlo con il metodo `push()`.

# --hints--

Dovresti usare il metodo `push()` per aggiungere `valueAtMiddle` all'array `pathToTarget`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /pathToTarget\s*\.\s*push\(\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
