---
id: 60f030d388cb74067cf291c3
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

見出しの下に、ユーザーが登録するよう促すために、次のテキストを段落要素内に使用してください。

```md
Please fill out this form with the required information
```

# --hints--

`p`の中に`body`要素を追加してください。

```js
assert.exists(document.querySelector('body > p'));
```

`p`の下に`h1`要素を追加してください。

```js
assert.exists(document.querySelector('h1 + p'));
```

`p`要素に`Please fill out this form with the required information`のテキストを設定してください。

```js
assert.equal(document.querySelector('p')?.innerText, 'Please fill out this form with the required information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
--fcc-editable-region--
<body>
  <h1>Registration Form</h1>
  
</body>
--fcc-editable-region--
</html>
```
