---
id: 60f02e7361b68405e27b62a5
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、登録フォームをデザインしてフォームのスタイルを学びます。

すべてのHTMLボイラープレート（`DOCTYPE`、`html`、`head`、`body`）は用意されています。

`body`の中に、`h1`要素を使ってテキスト`Registration Form`の見出しを追加し、フォームの文脈を示してください。

# --hints--

`h1`の中に`body`を追加してください。

```js
assert.exists(document.querySelector('body > h1'));
```

`h1`にテキスト`Registration Form`を設定してください。

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
