---
id: 60f1a5e2d2c23707a4f9a660
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

今度は、一部のブラウザーによって追加される`body`のデフォルトの`margin`を`0`に設定して、横スクロールバーをなくしましょう。

# --hints--

`margin`要素のセレクター内に`body`を追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.margin);
```

`margin`に`0`の値を設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
