---
id: 60f1922fcbd2410527b3bd89
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

前のレッスンでは、ビューポート単位の使い方を学びました。

`vh` 単位はビューポートの高さを表し、ビューポートの `height` の1％に相当することを覚えておいてください。これはビューポートの高さに対して相対的です。

プロジェクトをCSSで少し見栄えよくしましょう。まず、`body` に `width` を `100%` に、そして `height` を `100vh` に設定してください。

# --hints--

`body` 要素のセレクターを使うべきです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

`body` に `width` を `100%` に設定するべきです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

`body` に `height` を `100vh` に設定するべきです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
