---
id: 60f5dc35c07ac1078f140916
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

次のテキストを `label` 要素に追加してください：

- `Enter Your First Name:`
- `Enter Your Last Name:`
- `Enter Your Email:`
- `Create a New Password:`

# --hints--

最初の `label` にはテキスト `Enter Your First Name:` を入れてください。

```js
assert.match(document.querySelector('label')?.innerHTML, /Enter Your First Name:/i);
```

2番目の `label` にはテキスト `Enter Your Last Name:` を入れてください。

```js
assert.match(document.querySelector('fieldset > label:nth-child(2)')?.innerHTML, /Enter Your Last Name:/i);
```

3番目の `label` にはテキスト `Enter Your Email:` を入れてください。

```js
assert.match(document.querySelector('fieldset > label:nth-child(3)')?.innerHTML, /Enter Your Email:/i);
```

4番目の `label` にはテキスト `Create a New Password:` を入れてください。

```js
assert.match(document.querySelector('fieldset > label:nth-child(4)')?.innerHTML, /Create a New Password:/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label></label>
      <label></label>
      <label></label>
      <label></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
