---
id: 60f5cb8875ab6a0610f05071
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

フォームには3つの異なるセクションがあるので、`fieldset` 要素の中に3つの `form` 要素を追加してください。

# --hints--

3つの `fieldset` 要素を追加する必要があります。

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

3つの `fieldset` 要素は `form` の中にあるべきです。

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

3つの `fieldset` 要素は兄弟要素であるべきです。

```js
assert.exists(document.querySelector('fieldset + fieldset + fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
    
--fcc-editable-region--
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
