---
id: 60f805f813eaf2049bc2ceea
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

各`input`の中に`label`要素をネストしてください。各`input`は`label`のテキストの後に追加し、コロンの後にスペースを入れることを忘れないでください。

# --hints--

`input`要素に4つの`fieldset`要素を追加してください。

```js
assert.equal(document.querySelectorAll('fieldset input')?.length, 4);
```

`input`要素の中に`label`要素をネストしてください。

```js
assert.equal(document.querySelectorAll('label input')?.length, 4);
```

最初の`input`は`label`テキスト`Enter Your First Name:`の後に追加し、コロンの後にスペースを入れてください。

```js
const query = /^Enter Your First Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[0]?.innerHTML.trim(), query);
```

2番目の`input`は`label`テキスト`Enter Your Last Name:`の後に追加し、コロンの後にスペースを入れてください。

```js
const query = /^Enter Your Last Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[1]?.innerHTML.trim(), query);
```

3番目の`input`は`label`テキスト`Enter Your Email:`の後に追加し、コロンの後にスペースを入れてください。

```js
const query = /^Enter Your Email:\s+<input/
assert.match(document.querySelectorAll('label')?.[2]?.innerHTML.trim(), query);
```

4番目の`input`は`label`テキスト`Create a New Password:`の後に追加し、コロンの後にスペースを入れてください。

```js
const query = /^Create a New Password:\s+<input/
assert.match(document.querySelectorAll('label')?.[3]?.innerHTML.trim(), query);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label>Enter Your First Name:</label>
      <label>Enter Your Last Name:</label>
      <label>Enter Your Email:</label>
      <label>Create a New Password:</label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
