---
id: 60f81616cff80508badf9ad5
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

この時点で、フォームを提出できるはずです。しかし、あまり変化がないことに気づくかもしれません。

フォームをよりインタラクティブにするために、最初の`required`内の`input`要素に`fieldset`属性を追加してください。

これで、必須フィールドを入力せずにフォームを提出しようとすると、エラーメッセージが表示されます。

# --hints--

最初の`input`要素に`required`属性を付けてください。

```js
assert.isTrue(document.querySelector('input')?.required);
```

2番目の`input`要素に`required`属性を付けてください。

```js
assert.isTrue(document.querySelectorAll('input')?.[1]?.required);
```

3番目の`input`要素に`required`属性を付けてください。

```js
assert.isTrue(document.querySelectorAll('input')?.[2]?.required);
```

4番目の`input`要素に`required`属性を付けてください。

```js
assert.isTrue(document.querySelectorAll('input')?.[3]?.required);
```

タイプが`input`の`submit`には`required`属性を付けないでください。

```js
assert.isFalse(document.querySelector('input[type="submit"]')?.required);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
