---
id: 60fab4a123ce4b04526b082b
title: Step 24
challengeType: 0
dashedName: step-24
---

# --description--

現在、ユーザーはラジオ入力を選択せずにフォームを提出できます。以前は入力が必須であることを示すために`required`属性を使いましたが、この場合は両方の入力に`required`を追加するとユーザーに誤った情報が伝わるため、機能しません。

これを解決するには、2番目の`legend`内の`Account type (required)`要素の前に、テキスト`label`を持つ`fieldset`要素を追加して、必要な内容の文脈を提供します。そして、フォームが必須データを含んで提出されるように、`checked`入力に`Personal`属性を追加してください。

# --hints--

`input`要素は`type`の`radio`のままにしてください。

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-child(2) input[type="radio"]'), 2);
```

2番目の`legend`要素内に`fieldset`要素を追加してください。

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(2) > legend'), 1);
```

`Account type (required)`要素に`legend`テキストを追加してください。

```js
assert.equal(document.querySelector('fieldset:nth-of-type(2) > legend')?.textContent, 'Account type (required)');
```

`checked`入力に`Personal`属性を付与してください。

```js
assert.isTrue(document.querySelector('fieldset:nth-child(2) input[type="radio"]')?.hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
--fcc-editable-region--
      
--fcc-editable-region--
      <label><input type="radio" name="account-type" value="personal" /> Personal</label>
      <label><input type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
