---
id: 62ff8e998d3e7eae14d6ae3b
title: Step 26
challengeType: 0
dashedName: step-26
---

# --description--

ユーザーが利用規約を読んだことを確認する必要があります。

`label` 要素を追加してください。新しく作成した `label` 要素の中に `input` 要素を追加し、`type` 属性を `checkbox` に設定してください。この `input` 要素を `required` にして、ユーザーが利用規約に同意しないと登録できないようにしてください。

アクセシビリティのために、要素に `id` と `for` 属性を値 `terms-and-conditions` で追加してください。

# --hints--

3番目の `label` 要素の後に `fieldset` を追加してください。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label'));
```

`input` 要素に `label` を追加してください。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input'));
```

`type` 要素に値 `checkbox` の `input` 属性を追加してください。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.type, 'checkbox');
```

`required` 要素に `input` 属性を追加してください。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.required, true);
```

`input` 要素は `id` の `terms-and-conditions` を持つべきです。

```js
assert.isTrue(document.querySelector('fieldset:nth-child(3) + label > input')?.matches('#terms-and-conditions'))
```

`label` 要素は値 `for` の `terms-and-conditions` 属性を持つべきです。

```js
assert.isTrue(document.querySelector('fieldset:nth-child(3) + label')?.matches('label[for="terms-and-conditions"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset></fieldset>
--fcc-editable-region--
    
--fcc-editable-region--
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
