---
id: 62ff8b9dab5ac88e4d3d43a3
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

アクセシビリティのベストプラクティスに従い、`input` 要素と `label` 要素を `for` 属性で関連付けてください。

それぞれの `first-name` 属性には、`last-name`、`email`、`new-password`、`id` を値として使用してください。

# --hints--

最初の `input` 要素には、`id` の `first-name` を設定してください。

```js
assert.isTrue(document.querySelectorAll('input')?.[0]?.matches('#first-name'))
```

2番目の `input` 要素には、`id` の `last-name` を設定してください。

```js
assert.isTrue(document.querySelectorAll('input')?.[1]?.matches('#last-name'))
```

3番目の `input` 要素には、`id` の `email` を設定してください。

```js
assert.isTrue(document.querySelectorAll('input')?.[2]?.matches('#email'))
```

4番目の `input` 要素には、`id` の `new-password` を設定してください。

```js
assert.isTrue(document.querySelectorAll('input')?.[3]?.matches('#new-password'))
```

最初の `label` 要素には、値が `for` の `first-name` 属性を設定してください。

```js
assert.isTrue(document.querySelectorAll('label')?.[0]?.matches('label[for="first-name"]'))
```

2番目の `label` 要素には、値が `for` の `last-name` 属性を設定してください。

```js
assert.isTrue(document.querySelectorAll('label')?.[1]?.matches('label[for="last-name"]'))
```

3番目の `label` 要素には、値が `for` の `email` 属性を設定してください。

```js
assert.isTrue(document.querySelectorAll('label')?.[2]?.matches('label[for="email"]'))
```

4番目の `label` 要素には、値が `for` の `new-password` 属性を設定してください。

```js
assert.isTrue(document.querySelectorAll('label')?.[3]?.matches('label[for="new-password"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label>Enter Your First Name: <input /></label>
      <label>Enter Your Last Name: <input /></label>
      <label>Enter Your Email: <input /></label>
      <label>Create a New Password: <input /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
