---
id: 62ff8b9dab5ac88e4d3d43a3
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

アクセシビリティのベストプラクティスに従って、`for` 属性を使用して `input` 要素と `label` 要素を関連付けてください。

Use `first-name`, `last-name`, `email`, and `new-password` as values for the respective `id` attributes.

# --hints--

1 つ目の `input` 要素の `id` を `first-name` に設定する必要があります。

```js
assert.isTrue(document.querySelectorAll('input')?.[0]?.matches('#first-name'))
```

2 つ目の `input` 要素の `id` を `last-name` に設定する必要があります。

```js
assert.isTrue(document.querySelectorAll('input')?.[1]?.matches('#last-name'))
```

3 つ目の `input` 要素の `id` を `email` に設定する必要があります。

```js
assert.isTrue(document.querySelectorAll('input')?.[2]?.matches('#email'))
```

4 つ目の `input` 要素の `id` を `new-password` に設定する必要があります。

```js
assert.isTrue(document.querySelectorAll('input')?.[3]?.matches('#new-password'))
```

1 つ目の `label` 要素には `first-name` という値の `for` 属性が必要です。

```js
assert.isTrue(document.querySelectorAll('label')?.[0]?.matches('label[for="first-name"]'))
```

2 つ目の `label` 要素には `last-name` という値の `for` 属性が必要です。

```js
assert.isTrue(document.querySelectorAll('label')?.[1]?.matches('label[for="last-name"]'))
```

3 つ目の `label` 要素には `email` という値の `for` 属性が必要です。

```js
assert.isTrue(document.querySelectorAll('label')?.[2]?.matches('label[for="email"]'))
```

4 つ目の `label` 要素には `new-password` という値の `for` 属性が必要です。

```js
assert.isTrue(document.querySelectorAll('label')?.[3]?.matches('label[for="new-password"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label>Enter Your First Name: <input /></label>
      <label>Enter Your Last Name: <input /></label>
      <label>Enter Your Email: <input /></label>
      <label>Create a New Password: <input /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
