---
id: 6606927d010be4300a4e5330
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

يُنشئ التهيئة كائنًا فارغًا. لكن الأصناف يمكن أن تحتوي على طرق، وهي مثل الدوال المحلية لكل نسخة. داخل الصنف، تُعلن الطرق كما يلي:

```py
class ClassName:
    def method_name():
        pass
```

داخل صنف `Board`، استبدل `pass` بطريقة فارغة باسم `spam`.

# --hints--

يجب إعلان طريقة باسم `spam` داخل صنف `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
    
gameboard = Board()
--fcc-editable-region--
```
