---
id: 6606bf4561f8794a0d345919
title: الخطوة 27
challengeType: 20
dashedName: step-27
---

# --description--

حان الوقت لاختبار الطريقة `valid_in_row`. استدعِ `valid_in_row` على `gameboard`. مرر لها `0` و `8` كمعلمتين واطبع النتيجة.

لاحظ مرة أخرى كيف تُعرّف الطريقة بثلاث معلمات، لكنها تُستدعى بمعلمتين فقط لأن `self` تُمرر تلقائيًا كالكائن الذي تُستدعى عليه الطريقة.

# --hints--

يجب أن تطبع نتيجة استدعاء `valid_in_row(0, 8)` على `gameboard`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.valid_in_row(0, 8))")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--

--fcc-editable-region--
```
