---
id: 6606bc4e5535c0484990ccd5
title: الخطوة 24
challengeType: 20
dashedName: step-24
---

# --description--

بعد ذلك، ستعمل على طريقة تتحقق مما إذا كان يمكن إدخال رقم معين في صف محدد من لوحة السودوكو.

داخل صنف `Board`، أنشئ طريقة باسم `valid_in_row` وأعطها ثلاث معلمات: `self`، و`row`، و`num`. حيث تمثل `self` نسخة الصنف، و`row` و`num` هما فهرس الصف والرقم الذي سيتم التحقق منه على التوالي.


# --hints--

يجب أن تنشئ طريقة جديدة باسم `valid_in_row` داخل صنف `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_row")`)) })
```

يجب أن تحتوي طريقة `valid_in_row` على ثلاث معلمات: `self`، و`row`، و`num`، بهذا الترتيب.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_row").has_args("self, row, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
