---
id: 66069b992c1c5e3451f3deb0
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

الاستنساخ يُنشئ كائنًا فارغًا. الدالة `__init__` هي دالة خاصة تتيح لك استنساخ كائن بحالة مخصصة. عندما ينفذ الصنف دالة `__init__`، تُستدعى `__init__` تلقائيًا عند الاستنساخ.

داخل صنف `Board` الخاص بك، احذف طريقة `spam` واستبدلها بدالة `__init__` تتضمن معلمة `self`.

# --hints--

يجب ألا يكون لديك طريقة `spam` في صنف `Board` الخاص بك.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

يجب تعريف دالة `__init__` في صنف `Board` الخاص بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("__init__")`)) })
```

يجب أن تحتوي دالة `__init__` على معلمة `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
