---
id: 6606b224a69a293f98f8db8f
title: الخطوة 17
challengeType: 20
dashedName: step-17
---

# --description--

طريقة `.index()` ترفع استثناء `ValueError` عندما لا يُعثر على القيمة. لمنع توقف تنفيذ البرنامج، ضع هذا السطر من الشفرة داخل كتلة `try`. تُستخدم عبارة `try` لتغليف الشفرة التي قد ترفع استثناءً. أما بند `except` فيوفر شفرة بديلة للتنفيذ إذا حدث استثناء:

```py
try:
    <code>
except:
    <code>
```

ضع تعيين `col` داخل كتلة `try`. ثم أنشئ بند `except` واملأ جسمه بـ `pass`.

# --hints--

يجب وضع تعيين `col` داخل كتلة `try` وإنشاء بند `except`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_for_loops()[0].find_bodies()[0].is_equivalent("try:\\n  col = contents.index(0)\\nexcept:\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            col = contents.index(0)
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
