---
id: 66069167b3307b2f4067b22b
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

يُنشأ مثيل جديد من الصنف باستخدام تدوين الدالة، الذي يتضمن إلحاق زوج من الأقواس العادية باسم الصنف.

خارج تعريف الصنف، أنشئ مثيلًا من الصنف `Board` وقم بتعيينه إلى متغير اسمه `gameboard`.

# --hints--

يجب إعلان متغير `gameboard` خارج الصنف `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("gameboard")`)) })
```

يجب أن يحتوي متغير `gameboard` على القيمة `Board()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
--fcc-editable-region--
```
