---
id: 66068fb0bfddba2b7977eb60
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذا المشروع، ستتعلم عن الأصناف والكائنات من خلال بناء محلل لألغاز سودوكو.

في بايثون، الصنف هو مخطط لإنشاء الكائنات. الكائنات التي تُنشأ من الصنف هي نسخ من ذلك الصنف. يمكنك إنشاء صنف باستخدام هذه الصيغة النحوية:

```py
class ClassName:
    pass
```

حيث `class` هي الكلمة المفتاحية المطلوبة لتعريف الصنف و`ClassName` هو اسم الصنف، مكتوب وفقًا للاصطلاح في *PascalCase*.

ابدأ بإنشاء صنف `Board`.

# --hints--

يجب أن تنشئ صنفًا باسم `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
