---
id: 66069b0b36053733a2f012fe
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

احذف الآن استدعاء `spam` الخاص بك.

# --hints--

يجب ألا يحتوي شفرتك على `gameboard.spam()`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
gameboard.spam()
--fcc-editable-region--
```
