---
id: 6606933d6813a8308c962dd1
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

لكي تكون الطريقة طريقة مثيل، تحتاج إلى معلمة خاصة تُسمى `self` حسب الاصطلاح. هذه المعلمة هي مرجع لمثيل الصنف ويجب أن تكون دائمًا المعلمة الأولى.

أضف معلمة `self` إلى طريقتك `spam`.

# --hints--

يجب أن تحتوي طريقتك `spam` على معلمة `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam():
        pass
    
gameboard = Board()
--fcc-editable-region--
```
