---
id: 6606a3ccb1eea93a23c066bf
title: الخطوة 14
challengeType: 20
dashedName: step-14
---

# --description--

الآن ستعمل على طريقة تجد الخلايا الفارغة في لوحة السودوكو.

داخل صنف `Board`، أنشئ طريقة فارغة باسم `find_empty_cell` وأعطها معلمة `self`.

# --hints--

يجب أن تعرف طريقة باسم `find_empty_cell` داخل صنف `Board` الخاص بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("find_empty_cell")`)) })
```

يجب أن تحتوي طريقة `find_empty_cell` على معلمة `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        self.board = board
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
