---
id: 6606c14182435d4bab0de2ee
title: الخطوة 30
challengeType: 20
dashedName: step-30
---

# --description--

بعد ذلك، ستُنشئ طريقة تتحقق مما إذا كان يمكن إدخال رقم في عمود محدد من لوحة السودوكو عبر التحقق من أن الرقم غير موجود بالفعل في ذلك العمود.

داخل صنف `Board`، أنشئ طريقة باسم `valid_in_col` وأعطها ثلاث معلمات: `self`، و`col`، و`num`.
حيث `col` و`num` هما فهرس العمود والرقم المراد التحقق منه، على التوالي.

# --hints--

يجب أن تنشئ طريقة جديدة باسم `valid_in_col` داخل صنف `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_col")`)) })
```

يجب أن تحتوي طريقة `valid_in_col` على ثلاث معلمات: `self`، و`col`، و`num`، بهذا الترتيب.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_col").has_args("self, col, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
