---
id: 660699aabc59c532f2d556e5
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

لاستدعاء طريقة من كائن، تحتاج إلى استخدام ترميز النقطة:

```py
instance_name.method_name()
```

حيث `instance_name` هو الكائن أو النسخة، و`method_name` هي الطريقة التي تريد استدعاءها.

استدعِ الطريقة `spam` من الكائن `gameboard`.

# --hints--

يجب أن تستدعي الطريقة `spam` من الكائن `gameboard` باستخدام `gameboard.spam()`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
