---
id: 6606d8c323d6205890fbbd54
title: الخطوة 46
challengeType: 20
dashedName: step-46
---

# --description--

داخل صنف `Board`، أنشئ طريقة أخرى باسم `is_valid` وأعطها ثلاث معلمات: `self`، و`empty`، و`num`. حيث `empty` هو صف مرتب يمثل مؤشرات الصف والعمود لخلية فارغة و`num` هو الرقم الذي سيتم التحقق منه.

تتحقق هذه الطريقة مما إذا كان الرقم المعطى خيارًا صالحًا لخلية فارغة في لوحة السودوكو عبر التحقق من توافقه مع الصف والعمود والمربع 3x3 للخلية الفارغة المحددة.

# --hints--

يجب أن تنشئ طريقة جديدة باسم `is_valid` داخل صنف `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("is_valid")`)) })
```

يجب أن تحتوي طريقة `is_valid` على ثلاث معلمات: `self`، و`empty`، و`num`، بهذا الترتيب.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("is_valid").has_args("self, empty, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))

    def valid_in_square(self, row, col, num):
        row_start = (row // 3) * 3
        col_start = (col // 3) * 3
        for row_no in range(row_start, row_start + 3):
            for col_no in range(col_start, col_start + 3):
                if self.board[row_no][col_no] == num:
                    return False
        return True
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
