---
id: 660699119472f332798860ad
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

الآن، استبدل `pass` بـاستدعاء `print` ومرر له السلسلة النصية `'Spam!'`.

# --hints--

يجب حذف `pass` وطباعة السلسلة النصية `'Spam!'` داخل طريقة `spam`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").find_body().is_equivalent("print('Spam!')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        pass
    
gameboard = Board()
--fcc-editable-region--
```
