---
id: 6606a219f9efbf38ad496f67
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

يمكنك أيضًا استخدام ترميز النقطة للوصول إلى خاصية في نسخة من الصنف.

خارج صنف `Board`، بعد تهيئة كائن `gameboard`، استخدم `gameboard.board` للوصول إلى الخاصية `board` في كائن `gameboard` واطبع النتيجة على الشاشة.

# --hints--

يجب أن تطبع `gameboard.board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.board)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        self.board = board

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
