---
id: 6606a4641ec48b3a9fe8c2fc
title: الخطوة 15
challengeType: 20
dashedName: step-15
---

# --description--

تأخذ الدالة المدمجة `enumerate` قابلًا للتكرار كمعلمتها وتُرجع كائنًا من نوع enumerate يمكنك التكرار عليه.
توفر العدّ (الذي يبدأ افتراضيًا من الصفر) والقيمة من القابل للتكرار.

```py
iterable = ['a', 'b', 'c']
for i, j in enumerate(iterable):
    print(i, j)
```

الحلقة في المثال أعلاه ستُخرج الصفوف المرتبة `0, a`، `1, b`، و`2, c`.

داخل الطريقة `find_empty_cell`، استبدل `pass` بحلقة `for` تستخدم الدالة `enumerate()` للتكرار على كل صف في لوحة السودوكو. استخدم `row` كفهرس للصف الحالي و`contents` لعناصر الصف الحالي.


# --hints--

يجب أن تستبدل `pass` بحلقة `for` تتكرر على `enumerate(self.board)`. استخدم `row` و`contents` كمتغيرات للحلقة.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_body().is_equivalent("for row, contents in enumerate(self.board):\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
