---
id: 68e97fe79367ad7b5dd6c9cd
title: خطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

في هذه الخطوة أضف عنصر `h1` بالنص `XYZ Bookstore`.

# --hints--

يجب أن يكون لديك عنصر `h1`.

```js
assert.exists(document.querySelector('h1'));
```

يجب أن يكون نص عنصر `h1` هو `XYZ Bookstore`. تحقق جيدًا من الأخطاء الإملائية أو في الحروف الكبيرة.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'XYZ Bookstore');
```

يجب أن يكون عنصر `h1` داخل عنصر `body`.

```js
assert.match(code, /<body>[\s\S]*<h1>[\s\S]*<\/h1>[\s\S]*<\/body>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
