---
id: 68ebdbacdd3fa474132cc975
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

أسفل عنصر `h1`، أضف عنصر `p` يحتوي على هذا النص: `Browse our collection of amazing books!`.

# --hints--

يجب أن يكون لديك عنصر `p`.

```js
assert.exists(document.querySelector('p'));
```

يجب أن يكون نص عنصر `p` الخاص بك: `Browse our collection of amazing books!`.

```js
assert.equal(document.querySelector('p')?.innerText.trim(), 'Browse our collection of amazing books!');
```

يجب أن يكون عنصر `p` الخاص بك أسفل عنصر `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
