---
id: 68eca11e99e3c5c894ca9d69
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

يُستخدم عنصر `button` لإنشاء أزرار قابلة للنقر على صفحة الويب. الأزرار هي عناصر تفاعلية يمكن للمستخدمين النقر عليها لأداء إجراءات.

أضف عنصر `button` داخل العنصر الذي يحتوي على `class` بقيمة `card`، وأعطِ الزر خاصية `class` بقيمة `btn`، والنص `Buy Now`.

# --hints--

يجب أن يكون لديك عنصر `button` مسكن داخل العنصر الذي يحتوي على صنف `card`.

```js
assert.exists(document.querySelector('.card button'));
```

يجب أن يحتوي عنصر `button` على خاصية `class`.

```js
assert.isTrue(document.querySelector('.card button')?.hasAttribute('class'));
```

يجب أن تكون خاصية `class` في عنصر `button` بقيمة `btn`.

```js
assert.equal(document.querySelector('.card button')?.className, 'btn');
```

يجب أن يكون نص عنصر `button` هو `Buy Now`.

```js
assert.equal(document.querySelector('.card button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
