---
id: 68ec99e478211dc578699944
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

أسفل عنصر `h2` في أول عنصر يحمل الصنف `card`، أضف عنصر `p` بالنص التالي:

```md
This is an epic story of Sally and her dog Rex as they navigate through other worlds.
```

# --hints--

يجب أن يكون لديك عنصر `p` مسكن داخل العنصر الذي يحمل الصنف `card`.

```js
assert.exists(document.querySelector('.card p'));
```

يجب أن يكون نص عنصر `p` الخاص بك هو `This is an epic story of Sally and her dog Rex as they navigate through other worlds.`

```js
assert.equal(document.querySelector('.card p')?.innerText.trim(), "This is an epic story of Sally and her dog Rex as they navigate through other worlds.");
```

يجب أن يكون عنصر `p` الخاص بك أسفل عنصر `h2`.

```js
assert.exists(document.querySelector('.card h2 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
