---
id: 68eca5c0065c82d256d29ca3
title: الخطوة 17
challengeType: 0
dashedName: step-17
---

# --description--

داخل البطاقة الثانية، أضف عنصر `button` مع خاصية `class` مضبوطة على `btn` والنص `Buy Now`.

كلا عنصري `button` يشتركان الآن في نفس الـ `class`، مما يعني أنه يمكن تنسيقهما بشكل متناسق معًا.

# --hints--

يجب أن يكون لديك عنصر `button` مسكن داخل العنصر الثاني الذي يحمل صنف `card`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('button'));
```

يجب أن يحتوي عنصر `button` في بطاقتك الثانية على خاصية `class`.

```js
const cards = document.querySelectorAll('.card');
assert.isTrue(cards[1]?.querySelector('button')?.hasAttribute('class'));
```

يجب أن تكون خاصية `class` في عنصر `button` في بطاقتك الثانية بقيمة `btn`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.className, 'btn');
```

يجب أن يكون نص عنصر `button` في بطاقتك الثانية هو `Buy Now`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
