---
id: 68eca6364c5616d393389a30
title: الخطوة 18
challengeType: 0
dashedName: step-18
---

# --description--

تذكّر أن عنصر HTML يبدو هكذا:

```html
<element attribute="value">
    inner text
</element>
```

أسفل العنصر الذي يحمل الصنف `card-container`، أضف عنصر `p` جديدًا يحتوي على هذا النص:

```md
Review your selections and continue to checkout.
```

أسفل عنصر `p`، أنشئ عنصر `div` مع خاصية `class` مضبوطة على `btn-container`. ستجمع هذه الحاوية عناصر أزرار التنقل الخاصة بك.

# --hints--

يجب أن يكون عنصر `p` المضاف حديثًا أسفل العنصر الذي يحمل الصنف `card-container`.

```js
assert.exists(document.querySelector('.card-container + p'));
```

يجب أن يكون نص عنصر `p` الجديد هو `Review your selections and continue to checkout.`

```js
const cardContainer = document.querySelector('.card-container');
assert.equal(cardContainer?.nextElementSibling?.innerText.trim(), 'Review your selections and continue to checkout.');
```

يجب أن يكون لديك عنصر `div` مع `class` بقيمة `btn-container`.

```js
assert.exists(document.querySelector('div.btn-container'));
```

يجب أن يأتي عنصر `div` الذي يحمل `class` بقيمة `btn-container` بعد عنصر `p` الذي أُضيف حديثًا.

```js
assert.exists(document.querySelector('p + .btn-container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
