---
id: 68eca2a795b333ca5fee30a8
title: الخطوة 13
challengeType: 0
dashedName: step-13
---

# --description--

أنشئ الآن بطاقة كتاب ثانية. أضف عنصر `div` آخر مع خاصية `class` مضبوطة على `card`. لاحظ كيف يمكنك إعادة استخدام نفس اسم الصنف لعدة عناصر لتطبيق تنسيق متناسق.

# --hints--

يجب أن يكون لديك عنصران `div` مع خاصية `class` بقيمة `card`.

```js
assert.lengthOf(document.querySelectorAll('.card'), 2);
```

يجب أن يكون العنصر الثاني الذي يحمل صنف `card` داخل العنصر الذي يحمل صنف `card-container`.

```js
assert.lengthOf(document.querySelectorAll('.card-container .card'), 2);
```

يجب أن يكون العنصران اللذان يحملان صنف `card` شقيقين.

```js
assert.exists(document.querySelector('.card + .card'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
  --fcc-editable-region--
    
  --fcc-editable-region--
  </div>
</body>

</html>
```
