---
id: 68ec98b38d83a3c28dd30efe
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

داخل أول عنصر يحمل صنف `card`، أضف عنصر `h2` يحتوي على النص `Sally's SciFi Adventure`.

# --hints--

يجب أن يكون لديك عنصر `h2` مسكن داخل العنصر الذي يحمل صنف `card`.

```js
assert.exists(document.querySelector('.card h2'));
```

يجب أن يكون نص عنصر `h2` هو `Sally's SciFi Adventure`.

```js
assert.equal(document.querySelector('.card h2')?.innerText.trim(), "Sally's SciFi Adventure");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
