---
id: 68eca4c4b0f952cf09fabe09
title: الخطوة 15
challengeType: 0
dashedName: step-15
---

# --description--

داخل العنصر الثاني الذي يحمل صنف `card`، أضف عنصر `h2` يحتوي على النص `Dave's Cooking Adventure`.

# --hints--

يجب أن يكون لديك عنصر `h2` مسكن داخل العنصر الثاني الذي يحمل صنف `card`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2'));
```

يجب أن يحتوي عنصر `h2` داخل العنصر الثاني الذي يحمل صنف `card` على النص `Dave's Cooking Adventure`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('h2')?.innerText.trim(), "Dave's Cooking Adventure");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
