---
id: 68ec9332a9b5b2b32487bd00
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

تُضيف الخاصية `id` معرفًا فريدًا لعنصر HTML. يجب أن يكون كل `id` فريدًا داخل الصفحة ويُستخدم مرة واحدة فقط.

لا يمكن أن تحتوي قيم `id` على مسافات ويجب أن تحتوي فقط على حروف وأرقام وشرطات سفلية وشرطات.

إليك مثالًا:

```html
<p id="para">example paragraph</p>
```

أضف خاصية `id` إلى العنصر الذي يحمل الصنف `card` وحدد قيمتها إلى `sally-adventure-book`.

# --hints--

يجب أن يحتوي العنصر الذي يحمل الصنف `card` على خاصية `id`.

```js
assert.isTrue(document.querySelector('.card')?.hasAttribute('id'));
```

يجب أن يحتوي العنصر الذي يحمل الصنف `card` على خاصية `id` بقيمة `sally-adventure-book`.

```js
assert.equal(document.querySelector('.card')?.id, 'sally-adventure-book');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
  --fcc-editable-region--
    <div class="card">
  --fcc-editable-region--
      
    </div>
  </div>
</body>

</html>
```
