---
id: 68ec6cd5b7e8f5a8f7319e32
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

يُستخدم عنصر `div` كحاوية لتجميع عناصر HTML أخرى. ستستخدم عنصر `div` بشكل رئيسي عندما تريد تجميع عناصر HTML التي ستشارك مجموعة من أنماط CSS.

أسفل عنصر `p`، أضف عنصر `div`. سيكون هذا الـ `div` حاوية لبطاقات الكتب الخاصة بك.

**ملاحظة**: هذه الورشة لا تطبق CSS. الأصناف والعناصر المجمعة مفيدة لتنسيق CSS، لكن في هذه الورشة تُستخدم فقط لتنظيم وتجميع المحتوى. ستتعلم كيف يعمل التنسيق في وحدة لاحقة.

# --hints--

يجب أن يكون لديك عنصر `div`.

```js
assert.exists(document.querySelector('div'));
```

يجب أن يكون عنصر `div` الخاص بك أسفل عنصر `p`.

```js
assert.exists(document.querySelector('p + div'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
