---
id: 68ec6d9a315221aa31e54816
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

تُستخدم الخاصية `class` لتحديد عنصر أو أكثر من العناصر لأغراض التنسيق. على عكس الخاصية `id`، لا يجب أن تكون أسماء الأصناف فريدة: يمكن لعدة عناصر مشاركة نفس الصنف.

إليك مثالاً:

```html
<p class="example">example paragraph</p>
```

أضف خاصية `class` إلى عنصر `div` الخاص بك وضع قيمتها على `card-container`.

# --hints--

يجب أن يحتوي عنصر `div` الخاص بك على خاصية `class`.

```js
assert.isTrue(document.querySelector('div')?.hasAttribute('class'));
```

يجب أن يحتوي عنصر `div` الخاص بك على خاصية `class` بقيمة `card-container`.

```js
assert.equal(document.querySelector('div')?.className, 'card-container');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  <div>
--fcc-editable-region--
    
  </div>
</body>

</html>
```
