---
id: 69134b6cb43aa5f254950a10
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستبني صفحة مكتبة كتب بإنشاء بطاقات كتب تعرض معلومات عن كتب مختلفة. ستتدرب على تنظيم المحتوى باستخدام عناصر `div`، الأصناف، والمعرفات.

ابدأ صفحة المكتبة بإنشاء قالب جاهز لـ HTML.

أضف الإعلان `<!DOCTYPE html>` وعنصري `html` و`head`.

أضف خاصية `lang` إلى عنصر `html` واضبطها على `"en"`.

# --hints--

يجب أن يتضمن كودك إعلان `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

يجب أن يتضمن كودك وسمَي `<html>` افتتاحي وإغلاقي.

```js
assert.match(code,/<html\b[^>]*>[\s\S]*<\/html>/i);
```

يجب أن يحتوي عنصر `html` على خاصية `lang` مضبوطة على `"en"`.

```js
assert.match(code, /<html[^>]*\blang\s*=\s*["']en["'][^>]*>/i);
```

يجب أن يتضمن كودك وسمَي `<head>` افتتاحي وإغلاقي.

```js
assert.match(code,/<head>[\s\S]*<\/head>/i);
```

يجب أن يكون عنصر `head` مسكنًا داخل عنصر `html`.

```js
assert.match(code, /<html\b[^>]*>[\s\S]*<head>[\s\S]*<\/head>[\s\S]*<\/html>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
