---
id: 69134d14150d51f8066a3246
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

حسّن الآن هيكل وثيقة HTML الخاصة بك لضمان ترميز الصفحة بشكل صحيح.

داخل عنصر `head`، أضف العنصر `<meta charset="UTF-8">`.

أخيرًا، أضف عنصر `body` أسفل قسم `head`. هنا ستوضع كل محتويات الصفحة المرئية.

# --hints--

يجب أن تتضمن شفرتك عنصر `<meta charset="UTF-8">`.

```js
assert.exists(document.querySelector('head > meta[charset="utf-8"]'));
```

يجب أن تتضمن شفرتك وسمَي فتح وإغلاق `<body>`.

```js
assert.match(code,/<body>[\s\S]*<\/body>/i);
```

يجب أن يكون عنصر `body` أسفل عنصر `head`.

```js
assert.match(code, /<\/head>[\s\S]*<body>[\s\S]*<\/body>[\s\S]*<\/html>/i);
```

# --seed--
## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
  <title>XYZ Bookstore Page</title>
</head>

--fcc-editable-region--
</html>
```
