---
id: 691f85085787a8c4ca589fbb
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

أضف عنصر `title` داخل عنصر `head`.

عيّن عنوان الصفحة إلى `XYZ Bookstore Page`.

# --hints--

يجب أن تتضمن شفرتك وسم `<title>` الافتتاحي والختامي معًا.

```js
assert.match(code,/<title>[\s\S]*<\/title>/i);
```

يجب أن يكون عنوانك `XYZ Bookstore Page`.

```js
assert.equal(document.querySelector('title')?.innerText.trim(),'XYZ Bookstore Page');
```

يجب أن يكون عنصر `title` مسكنًا داخل عنصر `head`.

```js
assert.match(code, /<head>[\s\S]*<title>[\s\S]*<\/title>[\s\S]*<\/head>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
--fcc-editable-region-- 
</head>

</html>
```
