---
id: 68ed661e9f463c155bb677a1
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop üben Sie den Umgang mit CSS Flexbox, indem Sie eine Reihe farbiger Kästen gestalten.

Beginnen Sie damit, innerhalb des `body`-Elements ein `header`-Element hinzuzufügen. Verschachteln Sie ein `h1`-Element innerhalb des `header`-Elements mit dem Text `Colored Boxes Layout`.

# --hints--

Sie sollten ein `header`-Element innerhalb des `body`-Elements haben.

```js
const headerEl = document.querySelector('body > header');
assert.exists(headerEl);
```

Ihr `header`-Element sollte ein `h1`-Element enthalten.

```js
const h1El = document.querySelector('body > header > h1');
assert.exists(h1El);
```

Ihr `h1`-Element sollte den Text `Colored Boxes Layout` enthalten.

```js
const h1El = document.querySelector('body > header > h1');
assert.equal(h1El?.textContent.trim(), 'Colored Boxes Layout');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
</head>
<body>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```
