---
id: 68ef18bf52e0620eb8ae2694
title: Schritt 28
challengeType: 0
dashedName: step-28
---

# --description--

Jetzt überprüfen Sie den letzten Ausrichtungswert. Ändern Sie den Wert von `align-content` von `space-around` zu `space-evenly`.

Mit dem Wert `space-evenly` werden die Elemente gleichmäßig innerhalb des Ausrichtungscontainers verteilt. Der Abstand zwischen allen Elementen ist genau gleich, einschließlich des Abstands am Anfang und am Ende.

# --hints--

Ihr `.flex-container`-Selektor sollte die Eigenschaft `align-content` mit dem Wert `space-evenly` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('align-content'), 'space-evenly');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Colored Boxes</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <header>
    <h1>Colored Boxes Layout</h1>
  </header>
  <div class="flex-container">
    <div class="box">
      <h2>Box 1</h2>
      <p>Red</p>
    </div>
    <div class="box">
      <h2>Box 2</h2>
      <p>Orange</p>
    </div>
    <div class="box">
      <h2>Box 3</h2>
      <p>Yellow</p>
    </div>
    <div class="box">
      <h2>Box 4</h2>
      <p>Green</p>
    </div>
    <div class="box">
      <h2>Box 5</h2>
      <p>Blue</p>
    </div>
    <div class="box">
      <h2>Box 6</h2>
      <p>Indigo</p>
    </div>
  </div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
--fcc-editable-region--
  align-content: space-around;
--fcc-editable-region--
}

```
