---
id: 68ef0b3237fef991c92d9c19
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Fügen Sie nun ein `div`-Element mit der Klasse `flex-container` unter Ihrem `header` hinzu. Erstellen Sie dann sechs `div`-Elemente mit der Klasse `box` innerhalb des gerade erstellten `div`.

Denken Sie daran, dass ein Klassenattribut häufig verwendet wird, um auf einen Klassennamen in einem Stylesheet zu verweisen. In diesem Fall werden alle `div`-Elemente entsprechend der `.flex-container`- und `.box`-Style-Definition gleich gestylt.

# --hints--

Sie sollten ein `div`-Element unter dem `header`-Element haben.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.exists(divELBelowHeader);
```

Ihr `div`-Element unter dem `header`-Element sollte die Klasse `flex-container` haben.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.equal(divELBelowHeader.className, 'flex-container');
```

Sie sollten sechs `div`-Elemente verschachtelt innerhalb des `.flex-container`-Elements haben.

```js
const divElsInsideContainer = document.querySelectorAll('div.flex-container > div');
assert.lengthOf(divElsInsideContainer, 6);
```

Alle sechs Ihrer `div`-Elemente innerhalb des `.flex-container`-Elements sollten die Klasse `box` haben.

```js
const divElsInsideContainerWithBox = document.querySelectorAll('div.flex-container > div.box');
assert.lengthOf(divElsInsideContainerWithBox, 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```

```css
h1 {
	text-align: center;
	margin-bottom: 10px;
}
```
