---
id: 68f471d46fa193b002a5bfb8
title: Schritt 43
challengeType: 0
dashedName: step-43
---

# --description--

Erstellen Sie einen `.box6`-Selektor und fügen Sie die `background`-Eigenschaft mit dem Wert `#3f51b5` hinzu.

Fügen Sie dann die `order`-Eigenschaft mit dem Wert `5` zu Ihrem neuen Selektor hinzu. Sie werden sehen, dass sich das `.box6`-Element hinter das `.box5`-Element verschiebt.

Fügen Sie außerdem die `flex-grow`-Eigenschaft mit dem Wert `1` hinzu. Dadurch wächst Box 6, um den verbleibenden Platz im Container auszufüllen.

Damit ist der Workshop mit den farbigen Boxen abgeschlossen!

# --hints--

Sie sollten einen `.box6`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box6'));
```

Ihr `.box6`-Selektor sollte die Eigenschaft `background` mit dem Wert `#3f51b5` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box6')?.getPropVal('background'), 'rgb(63, 81, 181)');
```

Ihr `.box6`-Selektor sollte die Eigenschaft `order` mit dem Wert `5` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box6')?.getPropVal('order'), '5');
```

Ihr `.box6`-Selektor sollte die Eigenschaft `flex-grow` mit dem Wert `1` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box6')?.getPropVal('flex-grow'), '1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box box2">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box box3">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box box4">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box box5">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box box6">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

.box1 {
  background: #f16e79;
  order: 1; 
  flex-grow: 1;
}

.box2 {
  background: #f4a261;
  order: 0; 
}

.box3 {
  background: #ffd166;
  order: 2;
  flex-shrink: 3; 
}

.box4 {
  background: #4caf50;
  order: 3; 
}

.box5 {
  background: #457b9d;
  order: 4;
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box box2">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box box3">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box box4">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box box5">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box box6">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  height: 400px;
  align-content: space-evenly;
}

.box1 {
  background: #f16e79;
  order: 1; 
  flex-grow: 1; 
}

.box2 {
  background: #f4a261;
  order: 0; 
}

.box3 {
  background: #ffd166;
  order: 2; 
  flex-shrink: 3; 
}

.box4 {
  background: #4caf50;
  order: 3;
}

.box5 {
  background: #457b9d;
  order: 4; 
}

.box6 {
  background: #3f51b5;
  order: 5; 
  flex-grow: 1; 
}
```
