---
id: 68f46b758eb8ea7043abd73c
title: Schritt 33
challengeType: 0
dashedName: step-33
---

# --description--

Jetzt organisieren Sie die Boxen innerhalb Ihres Flexbox-Containers. Erstellen Sie zuerst einen `.box1`-Selektor und fügen Sie ihm die `background`-Eigenschaft mit dem Wert `#f16e79` hinzu.

Fügen Sie dann die Eigenschaft `order` mit dem Wert `1` hinzu. Sie werden sehen, wie die erste Box ans Ende verschoben wird. Das passiert, weil die `box`-Klasse bereits einen `order`-Wert von `0` hat, und Elemente mit höheren order-Werten später erscheinen.

Fügen Sie außerdem die Eigenschaft `flex-grow` mit dem Wert `1` hinzu. Dadurch wächst die Box, um den zusätzlichen Platz in ihrer Zeile auszufüllen.

# --hints--

Sie sollten einen `.box1`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box1'));
```

Ihr `.box1`-Selektor sollte die Eigenschaft `background` mit dem Wert `#f16e79` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('background'), 'rgb(241, 110, 121)');
```

Ihr `.box1`-Selektor sollte die Eigenschaft `order` mit dem Wert `1` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('order'), '1');
```

Ihr `.box1`-Selektor sollte die Eigenschaft `flex-grow` mit dem Wert `1` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('flex-grow'), '1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

--fcc-editable-region--

--fcc-editable-region--
```
