---
id: 68f93a8b58cbf330b4ba45c6
title: Schritt 22
challengeType: 0
dashedName: step-22
---

# --description--

Um die Gestaltung der `flex-container`-Klasse abzuschließen, fügen Sie die `padding`-Eigenschaft mit dem Wert `10px` und die `margin`-Eigenschaft mit dem Wert `20px auto` hinzu.

Denken Sie daran: `padding` fügt innerhalb des Elements Platz hinzu (zwischen dem Rahmen und dem Inhalt), während `margin` Platz außerhalb des Elements hinzufügt (zwischen dem Rahmen und anderen Elementen).

# --hints--

Ihr `.flex-container`-Selektor sollte die Eigenschaft `padding` mit dem Wert `10px` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('padding'), '10px');
```

Ihr `.flex-container`-Selektor sollte die Eigenschaft `margin` mit dem Wert `20px auto` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('margin'), '20px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
