---
id: 68f91df5776e693049fb8942
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Jetzt ist es an der Zeit, einige CSS-Regeln zur `styles.css`-Datei hinzuzufügen. Beginnen Sie damit, einen Selektor für das `h1`-Element zu erstellen.

Zentrieren Sie den Inhalt des `h1`-Elements, indem Sie seine `text-align`-Eigenschaft auf `center` setzen. Erstellen Sie außerdem eine `margin-bottom`-Eigenschaft mit dem Wert `10px`, um einen Außenabstand zwischen dem unteren Rand des `h1`-Elements und jedem HTML-Element darunter hinzuzufügen.

# --hints--

Sie sollten einen `h1`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

Ihr `h1`-Selektor sollte die `text-align`-Eigenschaft auf `center` setzen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('text-align'), 'center');
```

Ihr `h1`-Selektor sollte die `margin-bottom`-Eigenschaft auf `10px` setzen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('margin-bottom'), '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```

