---
id: 692f314d449329ba8bcd5d2e
title: Schritt 17
challengeType: 0
dashedName: step-17
---

# --description--

Die Eigenschaft `border-radius` rundet die Ecken des Rahmens eines Elements ab. Sie können angeben:

- Einen Wert, der auf alle vier Ecken angewendet wird.
- Zwei Werte, wobei der erste Wert für oben links/unten rechts und der zweite für oben rechts/unten links gilt.
- Drei Werte, die den Ecken oben links, oben rechts/unten links und unten rechts entsprechen.
- Vier Werte, die den Rahmenradius im Uhrzeigersinn beginnend an der oberen linken Ecke festlegen.

Fügen Sie nun die Eigenschaft `border-radius` mit dem Wert `5px` zum Selektor `.box` hinzu.

# --hints--

Ihr Selektor `.box` sollte die Eigenschaft `border-radius` mit dem Wert `5px` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('border-radius'), '5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Colored Boxes</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <header>
    <h1>Colored Boxes Layout</h1>
  </header>
  <div class="flex-container">
    <div class="box">
      <h2>Box 1</h2>
      <p>Red</p>
    </div>
    <div class="box">
      <h2>Box 2</h2>
      <p>Orange</p>
    </div>
    <div class="box">
      <h2>Box 3</h2>
      <p>Yellow</p>
    </div>
    <div class="box">
      <h2>Box 4</h2>
      <p>Green</p>
    </div>
    <div class="box">
      <h2>Box 5</h2>
      <p>Blue</p>
    </div>
    <div class="box">
      <h2>Box 6</h2>
      <p>Indigo</p>
    </div>
  </div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
