---
id: 692f22a2ad44ae887a24c8f9
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Fügen Sie in jedem `.box`-Element ein `h2`-Element mit dem Text `Box` gefolgt von einem Leerzeichen und einer fortlaufenden Nummer beginnend bei `1` hinzu.

# --hints--

Sie sollten ein `h2`-Element mit dem Text `Box 1` innerhalb des ersten `div` mit der Klasse `box` haben.

```js
const firstH2El = document.querySelector("div.box:nth-of-type(1) > h2");
assert.equal(firstH2El.textContent, "Box 1");
```

Sie sollten ein `h2`-Element mit dem Text `Box 2` innerhalb des zweiten `div` mit der Klasse `box` haben.

```js
const secondH2El = document.querySelector("div.box:nth-of-type(2) > h2");
assert.equal(secondH2El.textContent, "Box 2");
```

Sie sollten ein `h2`-Element mit dem Text `Box 3` innerhalb des dritten `div` mit der Klasse `box` haben.

```js
const thirdH2El = document.querySelector("div.box:nth-of-type(3) > h2");
assert.equal(thirdH2El.textContent, "Box 3");
```

Sie sollten ein `h2`-Element mit dem Text `Box 4` innerhalb des vierten `div` mit der Klasse `box` haben.

```js
const fourthH2El = document.querySelector("div.box:nth-of-type(4) > h2");
assert.equal(fourthH2El.textContent, "Box 4");
```

Sie sollten ein `h2`-Element mit dem Text `Box 5` innerhalb des fünften `div` mit der Klasse `box` haben.

```js
const fifthH2El = document.querySelector("div.box:nth-of-type(5) > h2");
assert.equal(fifthH2El.textContent, "Box 5");
```

Sie sollten ein `h2`-Element mit dem Text `Box 6` innerhalb des letzten `div` mit der Klasse `box` haben.

```js
const sixthH2El = document.querySelector("div.box:nth-of-type(6) > h2");
assert.equal(sixthH2El.textContent, "Box 6");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
		<div class="box">
			
		</div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
