---
id: 692f3001c9fc019f006f6b4e
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Die `align-items`-Eigenschaft richtet Flex-Elemente entlang der Querachse (senkrecht zur Hauptachse) aus.

Fügen Sie nun die `align-items`-Eigenschaft mit dem Wert `center` zum `.box`-Selektor hinzu. Da die `box`-Klasse `flex-direction: column` hat, ist die Querachse horizontal, sodass `center` die Elemente horizontal innerhalb des Containers zentriert.

# --hints--

Ihr `.box`-Selektor sollte die Eigenschaft `align-items` mit dem Wert `center` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('align-items'), 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
--fcc-editable-region--
  
--fcc-editable-region--  
}
```
