---
id: 692f22933ba1cb86cd75a70c
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Jetzt ist es Zeit, den Farbnamen für die Kästchen zu erstellen. Fügen Sie dazu unter jedem `p`-Element ein `h2`-Element mit diesen Farbnamen in der Reihenfolge hinzu: `Red`, `Orange`, `Yellow`, `Green`, `Blue` und `Indigo`. Eine Farbe für jedes Kästchen.

# --hints--

Sie sollten ein `p`-Element mit dem Text `Red` unter dem `h2`-Element innerhalb des ersten `.box`-Elements hinzufügen.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[0];
assert.equal(heading.textContent, "Red");
```

Sie sollten ein `p`-Element mit dem Text `Orange` unter dem `h2`-Element innerhalb des zweiten `.box`-Elements hinzufügen.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[1];
assert.equal(heading.textContent, "Orange");
```

Sie sollten ein `p`-Element mit dem Text `Yellow` unter dem `h2`-Element innerhalb des dritten `.box`-Elements hinzufügen.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[2];
assert.equal(heading.textContent, "Yellow");
```

Sie sollten ein `p`-Element mit dem Text `Green` unter dem `h2`-Element innerhalb des vierten `.box`-Elements hinzufügen.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[3];
assert.equal(heading.textContent, "Green");
```

Sie sollten ein `p`-Element mit dem Text `Blue` unter dem `h2`-Element innerhalb des fünften `.box`-Elements hinzufügen.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[4];
assert.equal(heading.textContent, "Blue");
```

Sie sollten ein `p`-Element mit dem Text `Indigo` unter dem `h2`-Element innerhalb des sechsten `.box`-Elements hinzufügen.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[5];
assert.equal(heading.textContent, "Indigo");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
    <div class="box">
      <h2>Box 1</h2>
			
    </div>
    <div class="box">
      <h2>Box 2</h2>
			
    </div>
    <div class="box">
      <h2>Box 3</h2>
			
    </div>
    <div class="box">
      <h2>Box 4</h2>
			
    </div>
    <div class="box">
      <h2>Box 5</h2>
			
    </div>
    <div class="box">
      <h2>Box 6</h2>
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
