---
id: 587d78af367417b2b2512b03
title: Jenga fomu ya utafiti
challengeType: 25
dashedName: build-a-survey-form
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kichwa cha ukurasa katika kipengele cha `h1` chenye `id` ya `title`.
1. Unapaswa kuwa na maelezo mafupi katika kipengele cha `p` chenye `id` ya `description`.
1. Unapaswa kuwa na kipengele cha `form` chenye `id` ya `survey-form`.
1. Ndani ya kipengele cha fomu unapaswa kuwa na sehemu ya ingizo ya `input` inayohitajika kuingiza jina lako ambayo ina `id` ya `name` na `type` ya `text`.
1. Ndani ya kipengele cha fomu unapaswa kuwa na sehemu ya ingizo ya `input` inayohitajika kuingiza barua pepe yako ambayo ina `id` ya `email`.
1. Ukiaingiza barua pepe isiyo na muundo sahihi, unapaswa kuona kosa la uthibitisho la HTML5.
1. Ndani ya kipengele cha fomu unapaswa kuwa na sehemu ya ingizo ya `input` ya kuingiza nambari ambayo ina `id` ya `number`.
1. Sehemu ya kuingiza nambari haipaswi kuruhusu kuingiza vitu visivyo nambari, iwe kwa kukuzuia kuviandika au kwa kuonyesha kosa la uthibitisho la HTML5 (kulingana na kivinjari chako).
1. Ukiaingiza nambari nje ya anuwai ya sehemu ya kuingiza nambari, ambayo imefafanuliwa na sifa za `min` na `max`, unapaswa kuona kosa la uthibitisho la HTML5.
1. Kwa sehemu za kuingiza jina, barua pepe, na nambari, unapaswa kuwa na vipengele vinavyolingana vya `label` katika fomu, vinavyoelezea kusudi la kila sehemu kwa vitambulisho vifuatavyo: `name-label`, `email-label`, na `number-label`.
1. Kwa sehemu za kuingiza jina, barua pepe, na nambari, unapaswa kuwa na maandishi ya kiwakilishi yanayotoa maelezo au maelekezo kwa kila sehemu.
1. Ndani ya kipengele cha fomu, unapaswa kuwa na kipengele cha orodha inayoshuka cha `select` chenye `id` ya `dropdown` na chenye chaguzi angalau mbili za kuchagua.
1. Ndani ya kipengele cha fomu, unaweza kuchagua chaguo kutoka kwa kundi la vitufe vya radio angalau viwili vilivyopangwa kwa kutumia sifa ya `name`.
1. Ndani ya kipengele cha fomu, unaweza kuchagua sehemu kadhaa kutoka kwa mfululizo wa visanduku vya kuchagua, kila kimoja kikiwa na sifa ya `value`.
1. Ndani ya kipengele cha fomu, unapaswa kuwa na sehemu ya `textarea` kwa maoni ya ziada.
1. Ndani ya kipengele cha fomu, unapaswa kuwa na kitufe chenye `id` ya `submit` ili kuwasilisha ingizo zote.

# --hints--

Unapaswa kuwa na kipengele cha `h1` chenye `id` ya `title`.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

`#title` yako haipaswi kuwa tupu.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Unapaswa kuwa na kipengele cha `p` chenye `id` ya `description`.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

`#description` yako haipaswi kuwa tupu.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Unapaswa kuwa na kipengele cha `form` chenye `id` ya `survey-form`.

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

Unapaswa kuwa na kipengele cha `input` chenye `id` ya `name`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

`#name` yako inapaswa kuwa na `type` ya `text`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

`#name` yako inapaswa kuhitaji ingizo.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

`#name` yako inapaswa kuwa mzao wa `#survey-form`.

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

Unapaswa kuwa na kipengele cha `input` chenye `id` ya `email`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

`#email` yako inapaswa kuwa na `type` ya `email`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

`#email` yako inapaswa kuhitaji ingizo.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

`#email` yako inapaswa kuwa mzao wa `#survey-form`.

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

Unapaswa kuwa na kipengele cha `input` chenye `id` ya `number`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

`#number` yako inapaswa kuwa mzao wa `#survey-form`.

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

`#number` yako inapaswa kuwa na `type` ya `number`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

`#number` yako inapaswa kuwa na sifa ya `min` yenye thamani ya nambari.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

`#number` yako inapaswa kuwa na sifa ya `max` yenye thamani ya nambari.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

Unapaswa kuwa na kipengele cha `label` chenye `id` ya `name-label`.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Unapaswa kuwa na kipengele cha `label` chenye `id` ya `email-label`.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Unapaswa kuwa na kipengele cha `label` chenye `id` ya `number-label`.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

`#name-label` yako inapaswa kuwa na maandishi yanayoelezea ingizo.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`#email-label` yako inapaswa kuwa na maandishi yanayoelezea ingizo.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`#number-label` yako inapaswa kuwa na maandishi yanayoelezea ingizo.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`#name-label` yako inapaswa kuwa mzao wa `#survey-form`.

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

`#email-label` yako inapaswa kuwa mzao wa `#survey-form`.

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

`#number-label` yako inapaswa kuwa mzao wa `#survey-form`.

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

`#name` yako inapaswa kuwa na sifa na thamani ya `placeholder`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

`#email` yako inapaswa kuwa na sifa na thamani ya `placeholder`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

`#number` yako inapaswa kuwa na sifa na thamani ya `placeholder`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Unapaswa kuwa na sehemu ya `select` yenye `id` ya `dropdown`.

```js
const el = document.getElementById('dropdown');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'SELECT');
```

`#dropdown` yako inapaswa kuwa na angalau vipengele viwili vya `option` vinavyoweza kuchaguliwa (visivyozimwa).

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert.isAtLeast(els.length, 2);
```

`#dropdown` yako inapaswa kuwa mzao wa `#survey-form`.

```js
const el = document.querySelector('#survey-form #dropdown');
assert.isNotNull(el);
```

Unapaswa kuwa na angalau vitufe viwili vya radio vya `input` vyenye `type` ya `radio`.

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

Unapaswa kuwa na angalau vitufe viwili vya radio vinavyokuwa mzao wa `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

Vitufe vyote vya radio vinapaswa kuwa na sifa na thamani ya `value`.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][value=""], input[type="radio"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Vitufe vyote vya radio vinapaswa kuwa na sifa na thamani ya `name`.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][name=""], input[type="radio"]:not([name])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Kila kundi la vitufe vya radio linapaswa kuwa na angalau vitufe viwili vya radio.

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  });
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert.fail();
});

assert.isAbove(groupKeys.length, 0);
```

Unapaswa kuwa na angalau visanduku viwili vya kuchagua vya `input` vyenye `type` ya `checkbox` vinavyokuwa mzao wa `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert.isAtLeast(els.length, 2);
```

Visanduku vyote vya kuchagua ndani ya `#survey-form` vinapaswa kuwa na sifa na thamani ya `value`.

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll(
  '#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Unapaswa kuwa na angalau kipengele kimoja cha `textarea` kinachokuwa mzao wa `#survey-form`.

```js
const el = document.querySelector('#survey-form textarea');
assert.isNotNull(el);
```

Unapaswa kuwa na kipengele cha `input` au `button` chenye `id` ya `submit`.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.isTrue(el.tagName === 'INPUT' || el.tagName === 'BUTTON');
```

`#submit` yako inapaswa kuwa na `type` ya `submit`.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.strictEqual(el.type, 'submit');
```

`#submit` yako inapaswa kuwa mzao wa `#survey-form`.

```js
const el = document.querySelector('#survey-form #submit');
assert.isNotNull(el);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Survey Form</title>
  </head>

  <body></body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="container">
      <header class="header">
        <h1 id="title" class="text-center">freeCodeCamp Survey Form</h1>
        <p id="description" class="description text-center">
          Thank you for taking the time to help us improve the platform
        </p>
      </header>
      <form id="survey-form">
        <div class="form-group">
          <label id="name-label" for="name">Name</label>
          <input
            type="text"
            name="name"
            id="name"
            class="form-control"
            placeholder="Enter your name"
            required
          />
        </div>
        <div class="form-group">
          <label id="email-label" for="email">Email</label>
          <input
            type="email"
            name="email"
            id="email"
            class="form-control"
            placeholder="Enter your Email"
            required
          />
        </div>
        <div class="form-group">
          <label id="number-label" for="number"
            >Age<span class="clue">(optional)</span></label
          >
          <input
            type="number"
            name="age"
            id="number"
            min="10"
            max="99"
            class="form-control"
            placeholder="Age"
          />
        </div>
        <div class="form-group">
          <p>Which option best describes your current role?</p>
          <select id="dropdown" name="role" class="form-control" required>
            <option disabled selected value>Select current role</option>
            <option value="student">Student</option>
            <option value="job">Full Time Job</option>
            <option value="learner">Full Time Learner</option>
            <option value="preferNo">Prefer not to say</option>
            <option value="other">Other</option>
          </select>
        </div>

        <div class="form-group">
          <p>Would you recommend freeCodeCamp to a friend?</p>
          <label>
            <input
              name="user-recommend"
              value="definitely"
              type="radio"
              class="input-radio"
              checked
            />Definitely</label
          >
          <label>
            <input
              name="user-recommend"
              value="maybe"
              type="radio"
              class="input-radio"
            />Maybe</label
          >

          <label
            ><input
              name="user-recommend"
              value="not-sure"
              type="radio"
              class="input-radio"
            />Not sure</label
          >
        </div>

        <div class="form-group">
          <p>What is your favorite feature of freeCodeCamp?</p>
          <select id="most-like" name="mostLike" class="form-control" required>
            <option disabled selected value>Select an option</option>
            <option value="challenges">Challenges</option>
            <option value="projects">Projects</option>
            <option value="community">Community</option>
            <option value="openSource">Open Source</option>
          </select>
        </div>

        <div class="form-group">
          <p>
            What would you like to see improved?
            <span class="clue">(Check all that apply)</span>
          </p>

          <label
            ><input
              name="prefer"
              value="front-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Front-end Projects</label
          >
          <label>
            <input
              name="prefer"
              value="back-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Back-end Projects</label
          >
          <label
            ><input
              name="prefer"
              value="data-visualization"
              type="checkbox"
              class="input-checkbox"
            />Data Visualization</label
          >
          <label
            ><input
              name="prefer"
              value="challenges"
              type="checkbox"
              class="input-checkbox"
            />Challenges</label
          >
          <label
            ><input
              name="prefer"
              value="open-source-community"
              type="checkbox"
              class="input-checkbox"
            />Open Source Community</label
          >
          <label
            ><input
              name="prefer"
              value="gitter-help-rooms"
              type="checkbox"
              class="input-checkbox"
            />Gitter help rooms</label
          >
          <label
            ><input
              name="prefer"
              value="videos"
              type="checkbox"
              class="input-checkbox"
            />Videos</label
          >
          <label
            ><input
              name="prefer"
              value="city-meetups"
              type="checkbox"
              class="input-checkbox"
            />City Meetups</label
          >
          <label
            ><input
              name="prefer"
              value="wiki"
              type="checkbox"
              class="input-checkbox"
            />Wiki</label
          >
          <label
            ><input
              name="prefer"
              value="forum"
              type="checkbox"
              class="input-checkbox"
            />Forum</label
          >
          <label
            ><input
              name="prefer"
              value="additional-courses"
              type="checkbox"
              class="input-checkbox"
            />Additional Courses</label
          >
        </div>

        <div class="form-group">
          <p>Any comments or suggestions?</p>
          <textarea
            id="comments"
            class="input-textarea"
            name="comment"
            placeholder="Enter your comment here..."
          ></textarea>
        </div>

        <div class="form-group">
          <button type="submit" id="submit" class="submit-button">
            Submit
          </button>
        </div>
      </form>
    </div>
  </body>
</html>
```
