---
id: 6842a6cd9836f0114a5b7a8a
title: 构建数字模式生成器
challengeType: 27
dashedName: build-a-number-pattern-generator
---

# --description--

在本实验中，你将通过创建一个小型应用来练习 Python 的基础知识，该应用用于创建一个数字模式。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该定义一个名为 `number_pattern` 的函数，该函数接受一个表示正整数的单个参数 `n`。
1. `number_pattern` 应该使用 `for` 循环。
1. `number_pattern(n)` 应该返回一个字串，包含从 1 到 `n`（包括 `n`）的所有整数，整数之间用空格分隔。例如，`number_pattern(4)` 应该返回字串 `1 2 3 4`。
1. 如果传递给函数的参数不是整数值，函数应返回 `Argument must be an integer value.`
1. 如果传递给函数的参数小于 1，函数应返回 `Argument must be an integer greater than 0.`

# --hints--

你应该有一个 `number_pattern` 函数。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('number_pattern')
    `)
})
```

`number_pattern` 函数应有一个名为 `n` 的参数。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('number_pattern').has_args('n')
    `)
})
```

`number_pattern(4)` 应该返回 `1 2 3 4`。

```js
({
    test: () => runPython(`
    assert number_pattern(4) == '1 2 3 4'
    `)
})
```

`number_pattern(12)` 应该返回 `1 2 3 4 5 6 7 8 9 10 11 12`。

```js
({
    test: () => runPython(`
    assert number_pattern(12) == '1 2 3 4 5 6 7 8 9 10 11 12'
    `)
})
```

`number_pattern` 应该为任何正整数返回一个以空格分隔的数字列表。

```js
({
    test: () => runPython(`
    from random import randint
    n = randint(1,65)
    ls = ' '.join(map(lambda n: str(n), list(range(1, n+1))))
    assert number_pattern(n) == ls

    n2 = randint(34,122)
    ls2 = ' '.join(map(lambda n: str(n), list(range(1, n2+1))))
    assert number_pattern(n2) == ls2
    
    `)
})
```

`number_pattern` 应该在传入非整数值时 `return` `Argument must be an integer value.`。

```js
({
    test: () => runPython(`
    string = 'Argument must be an integer value.'
    assert number_pattern('a') == string
    assert number_pattern(3.5) == string
    assert number_pattern(lambda n: n) == string
    `)
})
```

`number_pattern` 应该在传入非正整数时 `return` `Argument must be an integer greater than 0.`。

```js
({
    test: () => runPython(`
    from random import randint
    string = 'Argument must be an integer greater than 0.'
    assert number_pattern(0) == string
    assert number_pattern(-1185) == string
    assert number_pattern(-45) == string
    assert number_pattern(randint(-9999999, 0)) == string
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def number_pattern(n):
    result = "1"
    if isinstance(n, int):
        if n > 0:
            for i in range(2, n + 1):
                result += f" {i}"
            return result
        return "Argument must be an integer greater than 0."
    return "Argument must be an integer value."

```
