---
id: 68d275dd800f404d22a07564
title: реалізувати алгоритм пошуку у глибину
challengeType: 27
dashedName: lab-depth-first-search
---

# --description--

У цій лабораторній роботі ви реалізуєте алгоритм обходу графа, який називається <dfn>пошук у глибину</dfn>.

Якщо <dfn>пошук у ширину</dfn> досліджує ребра поступово, віддаляючись від початкового вузла, то <dfn>пошук у глибину</dfn> спочатку йде по шляху ребер якомога глибше.

Коли він досягає кінця шляху, пошук повертається назад до останнього вузла, у якого є неперевірений шлях, і продовжує пошук.

На відміну від пошуку у ширину, кожного разу, коли відвідується вузол, він не відвідує всіх його сусідів. Замість цього спочатку відвідує одного з сусідів і продовжує рухатися цим шляхом, поки не залишиться вузлів для відвідування на цьому шляху.

Для реалізації цього алгоритму варто використати стек (масив, у якому останній доданий елемент видаляється першим, відповідно до принципу <dfn>Last-In-First-Out</dfn>). Стек корисний у алгоритмах пошуку у глибину, бо коли ви додаєте сусідів у стек, ви хочете спочатку відвідати найновіше доданих сусідів і видалити їх зі стеку.

Простим виходом цього алгоритму є список вузлів, до яких можна дістатися з певного вузла. Тому також потрібно відслідковувати вузли, які ви відвідуєте.

**Мета**: Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути функція з назвою `dfs`.
1. Функція `dfs` має приймати два аргументи:
   - Неорієнтовану матрицю суміжності.
   - Мітку вузла, яка є числовим значенням вузла між `0` і `n - 1`, де `n` — загальна кількість вузлів у графі.
1. Функція `dfs` має реалізувати алгоритм пошуку у глибину, щоб вивести список усіх вузлів, до яких можна дістатися з переданого їй вузла.

# --hints--

У вас має бути функція з назвою `dfs`, яка приймає два аргументи.

```js
({ test: () => runPython(`
  import inspect
  assert inspect.isfunction(dfs)
  sig = inspect.signature(dfs)
  assert len(sig.parameters) == 2
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)` має повертати список з `1`, `2`, `3` і `0`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)
  assert isinstance(result, list)
  assert len(result) == 4
  assert set(result) == {1, 2, 3, 0}
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)` має повертати список з `1`, `2`, `3` і `0`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)
  assert isinstance(result, list)
  assert len(result) == 4
  assert set(result) == {3, 2, 1, 0}
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)` має повертати `[3]`.

```js
({ test: () => runPython(`
  assert dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3) == [3]
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)` має повертати список з `3` і `2`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)
  assert result == [3, 2] or result == [2, 3]
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)` має повертати список з `0` і `1`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)
  assert result == [0, 1] or result == [1, 0]
`) })
```

Функція `dfs` має повертати правильні результати.

```js
({ test: () => runPython(`
  import random

  def solve(g, s):
      v, q = set(), [s]
      while q:
          u = q.pop()
          if u not in v:
              v.add(u)
              q.extend([i for i, x in enumerate(g[u]) if x == 1 and i not in v])
      return v

  random.seed(0)

  for _ in range(10):
      n = 6
      mat = [[0]*n for _ in range(n)]
      for i in range(n):
          for j in range(i+1, n):
              if random.random() > 0.5:
                  mat[i][j] = mat[j][i] = 1

      start = random.randint(0, n-1)
      assert set(dfs(mat, start)) == solve(mat, start), "Random graph check failed"
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def dfs(graph, root):
    stack = []
    temp_v = None
    visited = []
    temp_v_neighbors = []
    stack.append(root)
    while stack:
        temp_v = stack.pop()
        if temp_v not in visited:
            visited.append(temp_v)
            temp_v_neighbors = graph[temp_v]
            for n, is_neighbor in enumerate(temp_v_neighbors):
                if is_neighbor == 1:
                    stack.append(n)
    return visited
```
