---
id: 66731cd027ef3acb155669f5
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will learn how to work with basic HTML elements like headings, paragraphs, and lists by building a cat photo app.

The first element you will learn about is the `h1` element. The `h1` element is a heading element and is used for the main heading of a webpage.

```html
<h1>This is a main heading</h1>
```

Add an `h1` element with the text of `CatPhotoApp` and watch the change in the browser preview.

When you are done, press the "Check Your Code" button to see if it's correct.

# --hints--

النص `CatPhotoApp` يجب أن يكون موجوداً في الكود. قد تحتاج إلى التحقق من الإملاء.

```js
assert.match(code, /catphotoapp/i);
```

يجب أن يحتوي العنصر `h1` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

يجب أن يحتوي العنصر `h1` على closing tag. Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Your `h1` element's text should be `CatPhotoApp`. You have either omitted the text, have a typo, or it is not between the `h1` element's opening and closing tags.

```js
assert.equal(document.querySelector('h1').innerText.toLowerCase(), 'catphotoapp');
```

You appear to be using a browser extension that is modifying the page. Be sure to turn off all browser extensions.

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
