---
id: 6690be42942f421648b119eb
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

You can add images to your website by using the `img` element. `img` elements have an opening tag without a closing tag. An element without a closing tag is known as a <dfn>void element</dfn>.

أضف عنصر `img` تحت عنصر `p`. في هذه المرحلة، لن تظهر أي صورة في المتصفح.

# --hints--

يجب أن يحتوي العنصر `img` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert.exists(document.querySelector('img'));
```

Your `img` element should not have a closing tag `</img>`.

```js
assert.notMatch(code, /<\/img\>/);
```

يجب أن يكون لديك عنصر `img` واحد فقط. قم بإزالة أي زيادات.

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

يجب أن يكون عنصر `img` الخاص بك تحت عنصر `p`. لقد وضعتهم بترتيب خاطئ.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```


