---
id: 6698444102ebd53d78cf66f8
title: الخطوة 22
challengeType: 0
dashedName: step-22
---

# --description--

After the unordered list, add a new image with a `src` attribute value set to:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

وقيمة السمة `alt` الخاصة بها إلى:

`A slice of lasagna on a plate.`

# --hints--

يجب أن يكون هناك عنصر `img` مباشرة بعد علامة الإغلاق `</ul>`.

```js
assert.equal(document.querySelectorAll('section')[1].lastElementChild.nodeName, 'IMG');
```

الصورة الجديدة لا تحتوي على سمة `alt`. تحقق من وجود مسافة بعد اسم opening tag و/أو أن هناك مسافات قبل جميع أسماء السمات.

```js
assert.isTrue(document.querySelectorAll('section')[1].lastElementChild.hasAttribute('alt'));
```

يجب أن تحتوي الصورة الجديدة على قيمة `alt` من `A slice of lasagna on a plate.` تأكد من أن قيمة السمة `alt` محاطة بعلامات اقتباس.

```js
assert.match(
  document.querySelectorAll('section')[1]
    .lastElementChild.getAttribute('alt')
    .replace(/\s+/g, ' '), /^A slice of lasagna on a plate\.?$/i
);
```

The new image does not have a `src` attribute. تيقن من وجود مسافة بعد اسم علامة الفتح و/أو أن هناك مسافات قبل جميع أسماء السمات.

```js
assert.isTrue(document.querySelectorAll('section')[1].lastElementChild.hasAttribute('src'));
```

The new image should have a `src` value of `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`. تأكد من أن قيمة السمة `src` محاطة بعلامات اقتباس.

```js
assert.equal(
  document.querySelectorAll('section')[1].lastElementChild.getAttribute('src'),
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

مع أنّك عيّنت لسمة `src` في عنصر img العنوان URL الصحيح، يوصى بأن تحيط دائما قيمة السمة بعلامات اقتباس.

```js
assert.notMatch(code, /\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```


