---
id: 66984e0cb4febb50c3794edf
title: الخطوة 33
challengeType: 0
dashedName: step-33
---

# --description--

The `footer` element is used to define a footer for a document or section. A footer typically contains information about the author of the document, copyright data, links to terms of use, contact information, and more.

بعد عنصر `main`، أضف عنصر `footer`.

# --hints--

إما أنك قمت بحذف عنصر `main` أو أنه يفتقد علامة فتح أو إغلاق.

```js
assert.exists(document.querySelector('main'));
assert.match(code, /<\/main>/);
```

يجب أن يحتوي العنصر `footer` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert.exists(document.querySelector('footer'));
```

يجب أن يحتوي العنصر `footer` على closing tag. Closing tags يجب ان تتبع بـ `/` مباشرة بعد رمز `<`.

```js
assert.match(code, /<\/footer\>/);
```

عنصر `footer` الخاص بك يجب أن يكون تحت عنصر `main` الخاص بك. لقد وضعته في مكان آخر.

```js
assert.equal(document.querySelector('main').nextElementSibling.nodeName, 'FOOTER');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
--fcc-editable-region--
    </main>

  </body>
--fcc-editable-region--
</html>
```


