---
id: 66984881117baa41da8bbde0
title: الخطوة 28
challengeType: 0
dashedName: step-28
---

# --description--

After the ordered list, add another `figure` element.

# --hints--

يجب أن يحتوي العنصر `figure` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('figure'), 2);
```

يجب أن يحتوي العنصر `figure` على closing tag. Closing tags لها `/` مباشرة بعد رمز `<`.

```js
assert.lengthOf(code.match(/<\/figure>/g), 2);
```

يجب أن يكون هناك عنصر `figure` أعلى closing tag عنصر `section` الثاني.

```js
assert.equal(document.querySelectorAll('main > section')[1].lastElementChild.nodeName, 'FIGURE');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
--fcc-editable-region--
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```


