---
id: 669842cee55b683aa13ee758
title: الخطوة 18
challengeType: 0
dashedName: step-18
---

# --description--

Within the second `section` element, add a new `h2` element with the text `Cat Lists`.

# --hints--

يجب أن يحتوي العنصر `section` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
assert.lengthOf(code.match(/<\/section>/g), 2);
```

يجب أن يحتوي العنصر `h2` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('h2'), 2);
```

يجب أن يحتوي العنصر `h2` على closing tag. Closing tags يجب ان تتبع بـ `/` مباشرة بعد رمز `<`.

```js
assert.lengthOf(code.match(/<\/h2\>/g), 2);
```

عنصر `h2` الثاني الخاص بك يجب أن يكون أعلى closing tag عنصر `section` الثاني. هو ليس في المكان الصحيح.

```js
const secondSection = document.querySelectorAll('section')[1];
assert.equal(secondSection.lastElementChild.nodeName, 'H2');
```

يجب أن يحتوي العنصر `h2` الثاني على النص `Cat Lists`. إما أنك حذفت النص أو لديك خطأ إملائي.

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase(), 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>

      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```


