---
id: 68f1eb9ad0a7575e4dbc31d5
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

بعد ذلك، أنشئ عنصر `video` أسفل عنصر `h1`. خلال الخطوات التالية، ستضيف الخاصيات اللازمة لجعل مشغل الفيديو يعمل.

# --hints--

يجب أن يكون لديك عنصر `video`.

```js
assert.exists(document.querySelector('video'));
```

يجب أن يكون عنصر `video` الجديد أسفل عنصر `h1`.

```js
assert.exists(document.querySelector('h1 + video'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
