---
id: 68f1f371fefd448b4f70b898
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

في درس سابق، تعلّمت عن الخصائص المختلفة المتاحة لعنصر `video`. الخاصية `width` تحدد عرض الفيديو بالبكسل.

أضف الخاصية `width` إلى عنصر `video` بقيمة `640`.

# --hints--

يجب أن يحتوي عنصر `video` على الخاصية `width` بالقيمة `640`.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('width'), '640');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video>
--fcc-editable-region--
  </video>
</body>
</html>
```
