---
id: 68f1f6f97af9749a42260a43
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

تُوفر الخاصية `controls` عناصر تحكم في التشغيل تشمل التشغيل، الإرجاع، والتحكم في مستوى الصوت لعنصر `video`.

الخاصية `controls` هي خاصية منطقية ولا تحتاج إلى قيمة.

أضف الخاصية `controls` إلى عنصر `video`.

الآن يجب أن ترى عنصر `video` معروضًا على الصفحة.

# --hints--

يجب أن يحتوي عنصر `video` على الخاصية `controls`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop>
--fcc-editable-region--
  </video>
</body>
</html>
```
