---
id: 68f1f8c248306da204655117
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

خاصية `muted` تُسكت الصوت عند التشغيل الأولي. إذا فعلت `controls`، سيتمكن المستخدم من إلغاء كتم الصوت. حذف خاصية `muted` سيشغل الصوت عند التشغيل الأولي.

خاصية `muted` هي خاصية قيمة منطقية ولا تحتاج إلى قيمة.

أضف خاصية `muted` إلى عنصر `video`.

# --hints--

يجب أن يحتوي عنصر `video` على خاصية `muted`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('muted'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop controls>
--fcc-editable-region--
  </video>
</body>
</html>
```
