---
id: 68f1fac259a4afabb478b960
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

ربما لاحظت أنك لم تربط بالفيديو الفعلي. ستفعل ذلك في المرحلة التالية. بالنسبة لأنواع ملفات الفيديو، هناك اختلافات في دعم المتصفحات. لتلبية هذا، يمكنك استخدام عناصر `source` داخل عنصر `video` وسيختار المتصفح أول `source` متوافق.

إليك مثال على عنصر `source`:

```html
<video controls width="250">
  <source src="src-url-goes-here" type="video-type-goes-here" />
</video>
```

عنصر `source` هو عنصر فارغ لذا لا يحتوي على وسم إغلاق.

أضف عنصر `source` داخل عنصر `video` الخاص بك.


# --hints--

يجب أن يكون لديك عنصر `source` داخل عنصر `video` الخاص بك.

```js
assert.exists(document.querySelector('video > source'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
