---
id: 68f1f993460f06a64c033061
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

خاصية `poster` هي صورة مصغرة للفيديو. تخيل مقاطع الفيديو التي تشاهدها على YouTube. تُعرض هذه الصورة أثناء تحميل الفيديو. إذا حُذفت الخاصية، تُعرض أول لقطة من الفيديو خلال مرحلة التحميل.

أضف الآن خاصية `poster` بالقيمة `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg` إلى عنصر `video` الخاص بك.

# --hints--

يجب أن يحتوي عنصر `video` على خاصية `poster` بالقيمة `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('poster'), 'https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
--fcc-editable-region--
    
--fcc-editable-region--
  >
  </video>
</body>
</html>
```
