---
id: 68f1f57723debf9383c75947
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

خاصية `loop` تعيد تشغيل الفيديو بمجرد انتهاء التشغيل. فكر في ميم إنترنت يعيد التشغيل بشكل متكرر. حذف خاصية `loop` يجعل الفيديو يُشغل مرة واحدة فقط.

خاصية `loop` هي خاصية قيمة منطقية ولا تحتاج إلى قيمة.

أضف خاصية `loop` إلى عنصر `video`.

# --hints--

يجب أن يحتوي عنصر `video` على خاصية `loop`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640">
--fcc-editable-region--
  </video>
</body>
</html>
```
