---
id: 68f26d52b13e0f5b6e2d2e09
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

نوع MIME شائع آخر هو نوع MIME الخاص بـ `video/webm`.

WebM هو تنسيق ملف وسائط سمعية وبصرية مفتوح المصدر طورته Google، ومصمم أساسًا لمحتوى الوسائط على الويب. يدعم ترميزات الفيديو مثل VP8 وVP9 وAV1، وترميزات الصوت مثل Vorbis وOpus، مما يجعله خيارًا شائعًا لعناصر الفيديو والصوت في HTML5.

أسفل أول عنصر `source` لديك، أضف عنصر `source` آخر وأعطه خاصية `src` بالقيمة `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm` وخاصية `type` بالقيمة `video/webm`.

# --hints--

يجب أن يكون لديك عنصر `source` ثاني.

```js
assert.exists(document.querySelector('video > source:nth-of-type(2)'));
```

يجب أن يحتوي عنصر `source` الثاني على خاصية `src`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('src'));
```

يجب أن تحتوي خاصية `src` في عنصر `source` الثاني على القيمة `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm');
```

يجب أن يحتوي عنصر `source` الثاني على خاصية `type`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('type'));
```

يجب أن تحتوي خاصية `type` في عنصر `source` الثاني على القيمة `video/webm`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('type'), 'video/webm');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
