---
id: 68f2701b2160c765f23b60f8
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

نوع MIME شائع آخر هو نوع MIME `video/ogg`.

Ogg هو تنسيق حاوية وسائط متعددة رقمية مصمم لتوفير بث فعال ومعالجة الوسائط الرقمية. تديره مؤسسة Xiph.Org وهو مجاني ومفتوح، غير مقيد ببراءات البرمجيات. اسمه مشتق من "ogging"، وهو مصطلح من لعبة الكمبيوتر Netrek.

أسفل عنصر `source` الثاني، أضف عنصر `source` ثالث وأعطه خاصية `src` بالقيمة `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg` وخاصية `type` بالقيمة `video/ogg`.

# --hints--

يجب أن يكون لديك عنصر `source` ثالث.

```js
assert.exists(document.querySelector('video > source:nth-of-type(3)'));
```

يجب أن يحتوي عنصر `source` الثالث على خاصية `src`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('src'));
```

يجب أن تحتوي خاصية `src` في عنصر `source` الثالث على القيمة `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg');
```

يجب أن يحتوي عنصر `source` الثالث على خاصية `type`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('type'));
```

يجب أن تحتوي خاصية `type` في عنصر `source` الثالث على القيمة `video/ogg`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('type'), 'video/ogg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
