---
id: 68f26a5f79e97c5032c10642
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

لقد استخدمت ملف فيديو بامتداد `mp4`، وتحتاج إلى إخبار المتصفح بذلك ليعرف كيف يقرأ الملف.

ستستخدم الخاصية `type` لتحديد نوع MIME كـ `video/mp4`.

MIME (امتدادات البريد الإلكتروني متعددة الأغراض) هو معيار لوصف المستندات بأشكال أخرى غير نص ASCII، مثل الصوت والفيديو والصور.

MP4، المعروف رسميًا باسم MPEG-4 الجزء 14، هو تنسيق حاوية وسائط متعددة رقمية. يُستخدم على نطاق واسع لتخزين الفيديو والصوت، لكنه يمكن أن يشمل أيضًا أنواع بيانات أخرى مثل الترجمة والصور الثابتة. صُممت ملفات MP4 للبث عبر الإنترنت وهي متوافقة مع العديد من الأجهزة والمنصات.

الآن، أضف الخاصية `type` والقيمة `video/mp4`.

# --hints--

يجب أن يحتوي عنصر `source` على خاصية `type`.

```js
const source = document.querySelector('source');
assert.isTrue(source?.hasAttribute('type'));
```

يجب أن يحتوي عنصر `source` على خاصية `type` بقيمة `video/mp4`.

```js
const source = document.querySelector('source');
assert.strictEqual(source?.getAttribute('type'), 'video/mp4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
--fcc-editable-region--
      
--fcc-editable-region--
    >
  </video>
</body>
</html>
```
