---
id: 68f27139bddf9f6bf2821acf
title: الخطوة 13
challengeType: 0
dashedName: step-13
---

# --description--

آخر عنصر `source` ستضيفه سيكون لنوع MIME `video/quicktime`.

QuickTime هو بنية وسائط متعددة قابلة للتوسيع أنشأتها Apple، تدعم تشغيل، وبث، وترميز، وتحويل مجموعة متنوعة من صيغ الوسائط الرقمية. ليس شائعًا مثل صيغة MP4، قد تحتاجه لدعم التطبيقات الموروثة.

أسفل عنصر `source` الثالث، أضف عنصر `source` رابع وأعطه خاصية `src` بالقيمة `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov` وخاصية `type` بالقيمة `video/quicktime`.

تهانينا! أكملت ورشة عمل مشغل الفيديو HTML.

# --hints--

يجب أن يكون لديك عنصر `source` رابع.

```js
assert.exists(document.querySelector('video > source:nth-of-type(4)'));
```

يجب أن يحتوي عنصر `source` الرابع على خاصية `src`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('src'));
```

يجب أن تحتوي خاصية `src` في عنصر `source` الرابع على القيمة `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov');
```

يجب أن يحتوي عنصر `source` الرابع على خاصية `type`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('type'));
```

يجب أن تحتوي خاصية `type` في عنصر `source` الرابع على القيمة `video/quicktime`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('type'), 'video/quicktime');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    controls
    loop
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
       src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov"
      type="video/quicktime"
    >
  </video>
</body>
</html>
```
