---
id: 69aaafeaa4037197c75587cd
title: Implemente uma fila
challengeType: 26
dashedName: implement-a-queue
---

# --description--

Neste laboratório, você vai implementar uma estrutura de dados de fila usando funções. Uma fila é uma estrutura de dados First-In-First-Out (FIFO) onde os elementos são adicionados ao final da fila e removidos do início.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de usuário:**

1. Você deve ter uma `enqueue` function que adiciona um elemento ao final da fila.
1. Você deve ter uma `dequeue` function que remove e retorna o elemento frontal da fila.
1. Você deve ter uma `front` function que retorna o elemento frontal da fila sem removê-lo.
1. Você deve ter uma função `size` que retorna o número de elementos na fila.
1. Você deve ter uma `isEmpty` função que retorna `true` se a fila estiver vazia e `false` caso contrário.

**Nota:** A maioria dos testes depende da função `enqueue`. Implemente essa primeiro, pois os testes para `dequeue`, `front`, `size` e `isEmpty` exigem adicionar elementos à fila.

# --hints--

Você deve ter uma `enqueue` função.

```js
assert.isFunction(enqueue);
```

A `enqueue` function deve adicionar um elemento ao final da fila.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

Você deve ter uma `dequeue` função.

```js
assert.isFunction(dequeue);
```

A `dequeue` function deve remover e retornar o elemento frontal da fila.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(dequeue(queue), 'first');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(size(queue), 1);
```

`dequeue` deve retornar `undefined` quando chamado em uma fila vazia.

```js
const queue = initQueue();
assert.isUndefined(dequeue(queue));
```

`enqueue` deve continuar a adicionar elementos ao final da fila após os dequeues.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(dequeue(queue), 'first');
enqueue(queue, 'third');
assert.strictEqual(front(queue), 'second');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(dequeue(queue), 'third');
```

Diferentes instâncias de fila não devem compartilhar `state`.

```js
const queueOne = initQueue();
const queueTwo = initQueue();
enqueue(queueOne, 'first');
enqueue(queueTwo, 'second');
assert.strictEqual(front(queueOne), 'first');
assert.strictEqual(front(queueTwo), 'second');
assert.strictEqual(size(queueOne), 1);
assert.strictEqual(size(queueTwo), 1);
```

Você deve ter uma `front` função.

```js
assert.isFunction(front);
```

A `front` function deve retornar o elemento frontal da fila sem removê-lo.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

`front` deve retornar `undefined` quando chamado em uma fila vazia.

```js
const queue = initQueue();
assert.isUndefined(front(queue));
```

Você deve ter uma `size` function.

```js
assert.isFunction(size);
```

A `size` function deve retornar o número correto de elementos na fila.

```js
const queue = initQueue();
assert.strictEqual(size(queue), 0);
enqueue(queue, 'first');
assert.strictEqual(size(queue), 1);
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(size(queue), 3);
dequeue(queue);
assert.strictEqual(size(queue), 2);
```

Você deve ter uma `isEmpty` função.

```js
assert.isFunction(isEmpty);
```

A `isEmpty` function deve retornar `true` para uma fila vazia.

```js
const queue = initQueue();
assert.isTrue(isEmpty(queue));
```

A `isEmpty` function deve retornar `false` para uma fila não vazia.

```js
const queue = initQueue();
enqueue(queue, 'element');
assert.isFalse(isEmpty(queue));
```

A `isEmpty` function deve retornar `true` depois que todos os elementos forem removidos da fila.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
dequeue(queue);
dequeue(queue);
assert.isTrue(isEmpty(queue));
```

# --seed--

## --seed-contents--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  
}

function dequeue(queue) {
  
}

function front(queue) {
  
}

function size(queue) {
  
}

function isEmpty(queue) {
  
}
```

# --solutions--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  queue.collection.push(element);
}

function dequeue(queue) {
  return queue.collection.shift();
}

function front(queue) {
  return queue.collection[0];
}

function size(queue) {
  return queue.collection.length;
}

function isEmpty(queue) {
  return queue.collection.length === 0;
}
```
