---
id: a0b5010f579e69b815e7c5d6
title: Eine intelligente Wortersetzungsfunktion erstellen
challengeType: 26
dashedName: build-a-smart-word-replacement-function
---

# --description--

In diesem Labor erstellen Sie eine Funktion, die eine Suche und Ersetzung im Satz anhand der übergebenen Argumente durchführt und dann den neuen Satz zurückgibt.

**Hinweis:** Bewahren Sie die Groß- oder Kleinschreibung des ersten Zeichens im ursprünglichen Wort, wenn Sie es ersetzen. Wenn Sie zum Beispiel das Wort `Book` durch das Wort `dog` ersetzen möchten, sollte es als `Dog` ersetzt werden.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und lassen Sie alle Tests erfolgreich durchlaufen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `myReplace` erstellen.
1. Die `myReplace`-Funktion sollte drei Argumente annehmen: einen String, ein zu ersetzendes Wort und das Wort, durch das es ersetzt werden soll.
1. Die `myReplace`-Funktion sollte den neuen String mit der Ersetzung zurückgeben.
1. Sie sollten die Groß- oder Kleinschreibung des ersten Zeichens im ursprünglichen Wort bewahren, wenn Sie es ersetzen.

# --hints--

Sie sollten eine Funktion mit dem Namen `myReplace` erstellen.

```js
assert.isFunction(myReplace);
```

`myReplace` sollte drei Argumente annehmen.

```js
assert.lengthOf(myReplace, 3);
```

`myReplace("Let us go to the store", "store", "mall")` sollte den String `Let us go to the mall` zurückgeben.

```js
assert.deepEqual(
  myReplace('Let us go to the store', 'store', 'mall'),
  'Let us go to the mall'
);
```

`myReplace("He is Sleeping on the couch", "Sleeping", "sitting")` sollte den String `He is Sitting on the couch` zurückgeben.

```js
assert.deepEqual(
  myReplace('He is Sleeping on the couch', 'Sleeping', 'sitting'),
  'He is Sitting on the couch'
);
```

`myReplace("I think we should look up there", "up", "Down")` sollte den String `I think we should look down there` zurückgeben.

```js
assert.deepEqual(
  myReplace('I think we should look up there', 'up', 'Down'),
  'I think we should look down there'
);
```

`myReplace("This has a spellngi error", "spellngi", "spelling")` sollte den String `This has a spelling error` zurückgeben.

```js
assert.deepEqual(
  myReplace('This has a spellngi error', 'spellngi', 'spelling'),
  'This has a spelling error'
);
```

`myReplace("His name is Tom", "Tom", "john")` sollte den String `His name is John` zurückgeben.

```js
assert.deepEqual(
  myReplace('His name is Tom', 'Tom', 'john'),
  'His name is John'
);
```

`myReplace("Let us get back to more Coding", "Coding", "algorithms")` sollte den String `Let us get back to more Algorithms` zurückgeben.

```js
assert.deepEqual(
  myReplace('Let us get back to more Coding', 'Coding', 'algorithms'),
  'Let us get back to more Algorithms'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function myReplace(str, before, after) {
  if (before.charAt(0) === before.charAt(0).toUpperCase()) {
    after = after.charAt(0).toUpperCase() + after.substring(1);
  } else {
    after = after.charAt(0).toLowerCase() + after.substring(1);
  }
  return str.replace(before, after);
}
```
