---
id: a5229172f011153519423690
title: Jenga kalkuleta jumla ya Fibonacci zisizo za kawaida
challengeType: 26
dashedName: build-an-odd-fibonacci-sum-calculator
---

# --description--

Katika maabara hii utajenga kalkuleta jumla ya Fibonacci zisizo za kawaida inayohesabu jumla ya nambari zote za Fibonacci zisizo za kawaida ambazo ni ndogo au sawa na nambari chanya iliyotolewa.

**Lengo**: Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kitendakazi cha `sumFibs` kinachokubali nambari kama hoja.
1. Kitendakazi cha `sumFibs` kinapaswa kurudisha jumla ya nambari zote za Fibonacci zisizo za kawaida ambazo ni ndogo au sawa na nambari iliyotolewa.
1. Mfululizo wa Fibonacci unaanza na `0` na `1`, na kila nambari inayofuata ni jumla ya nambari mbili zilizotangulia.
1. Ni nambari za Fibonacci zisizo za kawaida pekee ndizo zitakazoongezwa kwenye jumla.

# --hints--

Unapaswa kuwa na kitendakazi cha `sumFibs`.

```js
assert.isFunction(sumFibs);
```

`sumFibs(1)` inapaswa kurudisha nambari.

```js
assert.isNumber(sumFibs(1));
```

`sumFibs(1000)` inapaswa kurudisha `1785`.

```js
assert.strictEqual(sumFibs(1000), 1785);
```

`sumFibs(4000000)` inapaswa kurudisha `4613732`.

```js
assert.strictEqual(sumFibs(4000000), 4613732);
```

`sumFibs(4)` inapaswa kurudisha `5`.

```js
assert.strictEqual(sumFibs(4), 5);
```

`sumFibs(75024)` inapaswa kurudisha `60696`.

```js
assert.strictEqual(sumFibs(75024), 60696);
```

`sumFibs(75025)` inapaswa kurudisha `135721`.

```js
assert.strictEqual(sumFibs(75025), 135721);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumFibs(num) {
  let a = 0, b = 1;
  let sum = 0;

  while (b <= num) {
    if (b % 2 !== 0) {
      sum += b;
    }
    [a, b] = [b, a + b];
  }

  return sum;
}
```
