---
id: 68e6aff5325410118e76e7be
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті ви попрактикуєтеся у роботі з ARIA, створюючи доступний аудіоконтролер.

Для початку додайте елемент `h1` з текстом `Audio Controls`.

# --hints--

У вас має бути елемент `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Ваш елемент `h1` має бути всередині елемента `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Ваш елемент `h1` має містити текст `Audio Controls`.

```js
const h1Text = document.querySelector('h1')?.innerText.trim();
assert.strictEqual(h1Text, 'Audio Controls');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
