---
id: 68eac57fadcc662ce917e75e
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Далі додайте елемент `button` з атрибутом `type`, встановленим у значення `button`.

Текст елемента `button` має бути `Play`.

# --hints--

У вас має бути елемент `button`.

```js
assert.exists(document.querySelector('button'));
```

Ваш елемент `button` має мати атрибут `type`.

```js
const button = document.querySelector('button');
assert.isTrue(button?.hasAttribute('type'));
```

Ваш елемент `button` має мати атрибут `type`, встановлений у `button`.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.getAttribute('type'), 'button');
```

Ваш елемент `button` має містити текст `Play`.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.innerText.trim(), 'Play');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
