---
id: 68eb7f2397da251a40e04aa2
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Далі додайте елемент `div`, щоб містити всі елементи для керування гучністю.

Всередині елемента `div` вкладіть елемент `span` з атрибутом `id`, встановленим у `volume-label`.
Також текст для елемента `span` має бути `Volume`.

Пізніше на елемент `id` будуть посилатися атрибути ARIA, щоб допоміжні технології могли ідентифікувати цей текст як мітку для керування гучністю.

# --hints--

У вас має бути елемент `div` під вашим елементом `button`.

```js
assert.exists(document.querySelector('button + div'));
```

У вас має бути елемент `span` всередині вашого елемента `div`.

```js
assert.exists(document.querySelector('button + div > span'));
```

Ваш елемент `span` має мати атрибут `id`.

```js
const spanEl = document.querySelector('button + div > span');
assert.isTrue(spanEl?.hasAttribute('id'));
```

Ваш елемент `span` має мати атрибут `id` зі значенням `volume-label`.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.getAttribute('id'), 'volume-label');
```

Ваш елемент `span` має містити текст `Volume`.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.innerText.trim(), 'Volume');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

--fcc-editable-region--

--fcc-editable-region--

  </body>
</html>
```
