---
id: 68ec0a48c9d6db3594d768c3
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Тепер час додати ще один елемент `span`.

Цей елемент `span` має мати атрибут `id` зі значенням `volume-description`.
Також текст для цього елемента `span` має бути `Adjust the sound level`.

Подібно до попереднього елемента `span`, референс `id` цього `span` пізніше буде використаний у атрибутах ARIA, щоб допоміжні технології могли ідентифікувати цей текст як опис для регулятора гучності.

# --hints--

У вас має бути два елементи `span` всередині елемента `div`.

```js
const spans = document.querySelectorAll('div > span');
assert.lengthOf(spans, 2);
```

Ваш другий елемент `span` має мати атрибут `id`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.isTrue(secondSpanEl?.hasAttribute('id'));
```

Ваш другий елемент `span` має мати атрибут `id` зі значенням `volume-description`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.getAttribute('id'), 'volume-description');
```

Ваш другий елемент `span` має містити текст `Adjust the sound level`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.innerText.trim(), 'Adjust the sound level');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      --fcc-editable-region--
    
      --fcc-editable-region--
    </div>

  </body>
</html>
```
