---
id: 68ed488821845990c911a223
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Ваш доступний аудіоконтролер майже готовий, але залишився останній крок.

Додайте елемент `button` з атрибутом `type`, встановленим у `button`, і з текстом `Mute` під елементом `div`.

І з цією зміною ви завершили це практичне заняття!

# --hints--

У вас має бути елемент `button` під елементом `div`.

```js
const buttonEl = document.querySelector('div + button');
assert.exists(buttonEl);
```

Ваш елемент `button` має мати атрибут `type`.

```js
const buttonEl = document.querySelector('div + button');
assert.isTrue(buttonEl?.hasAttribute('type'));
```

Ваш елемент `button` має мати атрибут `type` зі значенням `button`.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.getAttribute('type'), 'button');
```

Ваш елемент `button` має містити текст `Mute`.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.innerText.trim(), 'Mute');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
           aria-labelledby="volume-label volume-description">
    </div>
  --fcc-editable-region--
    
  --fcc-editable-region--

  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
             aria-labelledby="volume-label volume-description">
    </div>

    <button type="button">Mute</button>

  </body>
</html>
```
