---
id: 68ed4204ab7edc66b6cb2383
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Як ви пам’ятаєте з попередніх уроків, атрибут `aria-labelledby` використовується, коли на сторінці вже є текст, який можна використати як мітку. У цьому випадку ви додали атрибут `id` до своїх елементів `span`, які тепер будуть використані тут.

Додайте атрибут `aria-labelledby` до елемента `input`, встановлений у значення `volume-label volume-description`.

# --hints--

Ваш елемент `input` має мати атрибут `aria-labelledby`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('aria-labelledby'));
```

Ваш елемент `input` має мати атрибут `aria-labelledby` зі значенням `volume-label volume-description`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
const ariaLabelledbyAttr = inputEl?.getAttribute('aria-labelledby');
assert.oneOf(ariaLabelledbyAttr, ['volume-label volume-description', 'volume-description volume-label']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    <input type="range" min="0" max="100" value="50">
    --fcc-editable-region--
    </div>

  </body>
</html>
```
