---
id: 68ed3a96e0ebcc36565b39e5
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Далі потрібно додати елемент `input` з атрибутом `type`, встановленим у `range`, щоб створити повзунок гучності. Встановіть атрибут `min` у `0`, атрибут `max` у `100` і атрибут `value` у `50`, щоб визначити рівень гучності за замовчуванням.

# --hints--

У вас має бути елемент `input` під другим елементом `span`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl);
```

Ваш елемент `input` має мати атрибут `type`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('type'));
```

Ваш елемент `input` має мати атрибут `type` зі значенням `range`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('type'), 'range');
```

Ваш елемент `input` має мати атрибут `min`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('min'));
```

Ваш елемент `input` має мати атрибут `min` зі значенням `0`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('min'), '0');
```

Ваш елемент `input` має мати атрибут `max`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('max'));
```

Ваш елемент `input` має мати атрибут `max` зі значенням `100`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('max'), '100');
```

Ваш елемент `input` має мати атрибут `value`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('value'));
```

Ваш елемент `input` має мати атрибут `value` зі значенням `50`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('value'), '50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </div>

  </body>
</html>
```
