---
id: 65f569725359e10d345bc52a
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

角度、速度、高さが与えられたときに、投射物の軌道を計算して描画するプログラムを作成します。

まず`math`をインポートしてください。このプロジェクトでよく使います。`math.radians`、`math.cos`、`math.sin`などの便利なメソッドが含まれています。

また、重力加速度の値や後で役立つ特殊記号の値を持つ変数を作成してください（これらはコピー＆ペーストで使ってください）。

```py
GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
```

# --hints--

`import math`を使って`math`モジュールをインポートしてください。

```js
({test: () => assert(runPython(`_Node(_code).has_import('import math')`))})
```

`GRAVITATIONAL_ACCELERATION`という変数を作成し、値を`9.81`にしてください。

```js
({test: () => runPython(`
assert _Node(_code).has_variable("GRAVITATIONAL_ACCELERATION"), "missing variable"
assert _Node(_code).find_variable("GRAVITATIONAL_ACCELERATION").is_equivalent('GRAVITATIONAL_ACCELERATION = 9.81'), "variable has wrong value"
`)})
```

`PROJECTILE`という変数を作成し、値を`"∙"`にしてください。

```js
({test: () => runPython(`
assert _Node(_code).has_variable("PROJECTILE"), "missing variable"
assert _Node(_code).find_variable("PROJECTILE").is_equivalent('PROJECTILE = "∙"'), "variable has wrong value"
`)})
```

`x_axis_tick`という変数を作成し、値を`"T"`にしてください。

```js
({test: () => runPython(`
assert _Node(_code).has_variable("x_axis_tick"), "missing variable"
assert _Node(_code).find_variable("x_axis_tick").is_equivalent('x_axis_tick = "T"'), "variable has wrong value"
`)})
```

`y_axis_tick`という変数を作成し、値を`"⊣"`にしてください。

```js
({test: () => runPython(`
assert _Node(_code).has_variable("y_axis_tick"), "missing variable"
assert _Node(_code).find_variable("y_axis_tick").is_equivalent('y_axis_tick = "⊣"'), "variable has wrong value"
`)})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
