---
id: 65fd4bd84561a14d3e720061
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

Pythonのクラスでは、クラス変数`__slots__`は特別な使い方があります。`__slots__`を宣言して文字列のシーケンスを代入すると、そのシーケンスに含まれる属性だけが作成可能になります。また、`__dict__`という特別な属性の作成を防ぎ、属性へのアクセスをより効率的にします。

クラス内で`__slots__`変数を使って、クラスが持つ属性を定義してください。`__slots__`に、`__init__`で定義された属性名のうち3つの文字列を含むタプルを代入してください。

# --hints--

`__slots__`には、`'__height'`、`'__speed'`、`'__angle'`を含むタプルを代入してください。

```js
({test: () => assert(runPython(`
from itertools import permutations
slots = _Node(_code).find_class("Projectile").find_body().find_variable("__slots__")
perms = permutations(("__height", "__speed", "__angle"))
any(slots.is_equivalent(f'__slots__ = {perm}') for perm in perms)
`))})

```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--
class Projectile:

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
--fcc-editable-region--
```
