---
id: 65fd4b3bedee044b4b957d33
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

3つの引数、すなわち投射物の投げ始めの速度、始めの高さ、始めの角度の順でクラスを初期化する`Projectile`メソッドを持つ`__init__`クラスを作成してください。

`__init__`メソッド内で、これらの引数を3つのプライベート属性`__speed`、`__height`、`__angle`に代入してください。角度は度数で与えられますが、内部的にはラジアンで保存する必要があります。そのため、角度を度からラジアンに変換するために`math.radians()`関数を使い、`__angle`に代入してください。

属性名の前にアンダースコア2つを付けると、Pythonでは名前修飾が行われます。これは、属性がクラスの外部から直接アクセスできず、必要に応じて`ball._Projectile__height`のような修飾された名前でアクセスしなければならないことを意味し、これらは内部用であることを示しています。

# --hints--

クラスを作成するには`class Projectile:`を使ってください。

```js
({test: () => assert(runPython(`_Node(_code).has_class("Projectile")`))})

```

`__init__`クラスの中に`Projectile`メソッドがあるはずです。

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_body().has_function("__init__")`))})
```

新しい`Projectile`クラスのインスタンスは、3つの属性の値が正しく初期化されている必要があります。

```js
({test: () => assert(runPython(`
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
p = Projectile(20, 21, 22)
p._Projectile__height == 21 and p._Projectile__speed == 20 and p._Projectile__angle == math.radians(22)
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--

--fcc-editable-region--
```
