---
id: 660fcf3e1b9bb056b2edb567
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

ここまでで文字列の表現を作成する準備ができました。
まず、`Projectile` クラスのインスタンスを作成してください。クラス定義の下に、`ball` 変数を作成し、`Projectile` を引数として使い、`10, 3, 45` を呼び出して代入してください。

クラス内に `__str__` メソッドを定義し、インスタンスに適切な文字列の表現を与え、次の形式の文字列を返すようにしてください。

```py

Projectile details:
speed: 10 m/s
height: 3 m
angle: 45°
displacement: 12.6 m

```

文字列は改行文字で始まり改行文字で終わる必要があります。角度は度数の整数で表し、変位は小数点以下1桁で出力してください。

角度をラジアンから度に変換するのに役立つ `math.degrees` が見つかるでしょう。

準備ができたら、`print(ball)` して関数をテストしてください。

# --hints--

`ball` 変数があるはずです。

```js
({test: () => assert(runPython(`_Node(_code).has_variable('ball')`))})
```

`ball` 変数は `Projectile(10, 3, 45)` の値を持つはずです。

```js
({test: () => assert(runPython(`_Node(_code).find_variable('ball').is_equivalent('ball = Projectile(10, 3, 45)')`))})
```

`Projectile(45, 45, 45)` の文字列表現は正しいはずです。

```js
({test: () => assert(runPython(
`
ball = Projectile(45, 45, 45)
str(ball) == """
Projectile details:
speed: 45 m/s
height: 45 m
angle: 45°
displacement: 244.4 m
"""`
))})
```

他のインスタンスでも文字列表現は正しいはずです。

```js
({test: () => assert(runPython(`p = Projectile(10, 10, 10)
str(
    p
) == """
Projectile details:
speed: 10 m/s
height: 10 m
angle: 10°
displacement: 15.9 m
"""`))})
```

`print(ball)` の呼び出しがあるはずです。

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls('print')
    assert any(c.is_equivalent('print(ball)') for c in calls)
    `)
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)


    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
--fcc-editable-region--

--fcc-editable-region--
        
```
