---
id: 6616ec0bf9ee8548ce6b0f08
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

軌道の座標を計算する作業に取りかかる時間です！`Projectile` クラスに `__calculate_y_coordinate` という名前のメソッドを作成してください。このメソッドは `self` に加えて `x` のパラメータを持つ必要があります。

\\[ y = y_0 + x \tan\theta - \frac{g x^2}{2 v_0^2 \cos^2\theta} \\]

上記は、開始角度 $\theta$、速度 $v_0$、高さ $y_0$ を持つ任意の水平位置 $x$ に対して垂直位置 $y$ を計算するための式です。

`math.tan()` と `math.cos()` を使う必要があり、`x ** y` は $x^y$ を表す書き方であること、$g$ の値は変数 `GRAVITATIONAL_ACCELERATION` に入っていることを覚えておいてください。

このメソッドを実装して、$y$ 座標を返すようにしてください。

メソッドをテストしたい場合は、クラスの外から `ball._Projectile__calculate_y_coordinate()` が宣言された後に、数値を引数として `ball` を呼び出すことができます。

# --hints--

あなたの `Projectile` クラスには `__calculate_y_coordinate` という名前のメソッドがあるはずです。

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').has_function('__calculate_y_coordinate')`))}})
```

`__calculate_y_coordinate` メソッドは `self, x` を引数に持つべきです。

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').find_function('__calculate_y_coordinate').has_args('self, x')`))}})
```

メソッドは正しい値を返すべきです。`Projectile(12, 12, 12)` に対して、$x$ 座標が 0、1、2、3 のとき、$y$ 座標はおよそ 12、12.18、12.28、12.32 となるはずです。

```js
({test: () => {
    assert(runPython(`
    a = Projectile(12, 12, 12)
    all([round(a._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(0, 12.0), (1, 12.18), (2, 12.28), (3, 12.32)]])
    `))
}})
```

メソッドは正しい値を返すべきです。`Projectile(45, 12, 22)` に対して、$x$ 座標が 6、9、10、12 のとき、$y$ 座標はおよそ 14.32、15.41、15.76、16.44 となるはずです。

```js
({test: () => {
    assert(runPython(`
    b = Projectile(45, 12, 22)
    all([round(b._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(6, 14.32), (9, 15.41), (10, 15.76), (12, 16.44)]])
    `))
}})
```

`__calculate_y_coordinate` メソッドは結果を丸めてはいけません。

```js
({test: () => assert.isFalse(runPython(`Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5) == round(Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5), 2)`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    --fcc-editable-region--
    
    --fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
   
```
