---
id: 662f9f65c50bd21d05ffdee9
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

`calculate_all_coordinates` という名前のメソッドを定義してください。このメソッドは `0` から変位を切り上げた値（含まない）までのすべての \\(x\\) 値の座標を計算し、それらをタプルのリスト `(x, y)` として返します。

次に、`calculate_all_coordinates` インスタンスの `ball` メソッドを呼び出し、その出力を `coordinates` という新しい変数に代入してください。

数値をその数以上の最小の整数に切り上げるには `math.ceil()` を使うことができます。

# --hints--

引数が1つの `calculate_all_coordinates` メソッドを定義してください。引数は `self` です。

```js
({test: () => assert(runPython(`
c = _Node(_code).find_class('Projectile')
c.has_function('calculate_all_coordinates') and c.find_function('calculate_all_coordinates').has_args('self')
`))})
```

`calculate_all_coordinates` の `ball` メソッドを呼び出し、その結果を `coordinates` という変数に代入してください。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('coordinates'), "coordinates variable is missing"

        assert _Node(_code).find_variable('coordinates').is_equivalent('coordinates = ball.calculate_all_coordinates()'), "wrong value assigned to coordinates variable"
    `)
})
```

`calculate_all_coordinates` メソッドは、すべての `(x, y)` 座標を含むリストを返す必要があります。

```js
({test: () => assert(runPython(`
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
ball = Projectile(12, 13, 14)
expected = [(x, ball._Projectile__calculate_y_coordinate(x)) for x in range(math.ceil(ball._Projectile__calculate_displacement()))]
actual = ball.calculate_all_coordinates()
expected == actual
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x**2 / (
                2 * self.__speed**2 * math.cos(self.__angle)**2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate

    --fcc-editable-region--
    
        
        
ball = Projectile(10, 3, 45)
print(ball)
   --fcc-editable-region--
```
