---
id: 662fa453ec8033322b3ab2e3
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

カプセル化とゲッターについて説明する時間です。先頭にダブルアンダースコアを付けて、3つのインスタンス属性をプライベートにしました。これらの属性は慣例としてプライベートと呼ばれていますが、外部からアクセスすることは可能であっても、そうしないことが合意されています。

ゲッターは外部から値を取得するために使われます。ゲッターを定義するには、目的の属性の値を返すメソッドを定義し、`@property`デコレーターを付けます。

```py
class Nest:
    ...
    @property
    def number_of_eggs(self):
        return self.__number_of_eggs
```

このデコレーターはメソッドをプロパティに変えます。つまり、通常のメソッドのように呼び出すのではなく、属性のように使われます。

```py
n = Nest()
print(n.number_of_eggs)
```

上の例では、プライベート属性`__number of eggs`が`number_of_eggs`の`n`プロパティを通じてアクセスされています。

`speed`という名前のゲッターを作成し、プライベート属性`__speed`の値を返すようにしてください。

# --hints--

`speed`メソッドを定義してください。

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').has_function('speed')
`))})
```

`speed`メソッドはパラメータを1つ持ち、`self`である必要があります。

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_args('self')
`))})
```

`speed`メソッドには`@property`デコレーターを付けてください。

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_decorators('property')
`))})
```

`speed`メソッドはプライベート属性`__speed`の値を返す必要があります。

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.speed == 22
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

--fcc-editable-region--
    
--fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()
   
```
