---
id: 6633b26c4f2d603072ef0818
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

クラスに対して `__repr__` 特殊メソッドを使って表現を与えるのは良い習慣です。`__str__` メソッドはユーザーに優しい読みやすい文字列の表現を返しますが、`__repr__` はプログラマー向けです。多くの場合、`__repr__` はオブジェクトを再作成するために使える文字列を提供します。

オブジェクトをインスタンス化するために必要な文字列を返す `__repr__` メソッドを書いてください。

これは `repr` 関数を呼び出してクラスのインスタンスを渡すことでテストできます。

# --hints--

`Projectile` クラスは `__repr__` メソッドを持つべきです。

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_class('Projectile').has_function('__repr__')`
    ))
})
```

インスタンス `Projectile(12, 12, 12)` の表現は `'Projectile(12, 12, 12)'` と等しくなるはずです。

```js
({
    test: () => assert(runPython(
        `repr(Projectile(12, 12, 12)) == 'Projectile(12, 12, 12)' or \\
        repr(Projectile(12, 12, 12)) == 'Projectile(12, 12, 12.0)'`
    ))
})
```

インスタンス `Projectile(21, 22, 45)` の表現は `'Projectile(21, 22, 45)'` と等しくなるはずです。

```js
({
    test: () => assert(runPython(
        `repr(Projectile(21, 22, 45)) == 'Projectile(21, 22, 45)' or \\
        repr(Projectile(21, 22, 45)) == 'Projectile(21, 22, 45.0)'`
    ))
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.speed} m/s
height: {self.height} m
angle: {self.angle}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

    @property
    def height(self):
        return self.__height

    @property
    def angle(self):
        return round(math.degrees(self.__angle))

    @property
    def speed(self):
        return self.__speed
    
    @height.setter
    def height(self, n):
        self.__height = n

    @angle.setter
    def angle(self, n):
        self.__angle = math.radians(n)

    @speed.setter
    def speed(self, s):
       self.__speed = s

--fcc-editable-region--
    
--fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()

   
```
