---
id: 64b171849f925b0773aa434c
title: 步驟4
challengeType: 20
dashedName: step-4
---

# --description--

創建一個從 1 到 3 的數字列表，並將其分配到變量 `xs`。

# --hints--

`xs` 應該是列表 `[1,2,3]`。

```js
({ test: () => assert.deepEqual(Array.from(__userGlobals.get("xs")), [1,2,3])})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
xs = [1,2,3]
```
