---
id: 6720b3191bbafd269bae0ae8
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`div` 요소를 만들고 `id`인 `greeting-card`과 `class`인 `card`을 가지게 하세요.

`div` 요소 안에 `h1` 텍스트가 있는 `Happy Birthday!`를 추가하세요. 그런 다음 `class`라는 `message`을 가진 단락 요소와 `Wishing you all the happiness and joy on your special day!` 텍스트를 추가하세요.

# --hints--

`div` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("div"));
```

`div` 요소는 `id`인 `greeting-card`을 가져야 합니다.

```js
assert.strictEqual(document.querySelector("div").id, "greeting-card");
```

`div` 요소는 `class`인 `card`을 가져야 합니다.

```js
assert.strictEqual(document.querySelector("#greeting-card").className, "card");
```

`h1` 요소 안에 `div` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("div > h1"));
```

`h1` 요소는 `Happy Birthday!` 텍스트를 포함해야 합니다.

```js
assert.strictEqual(document.querySelector("div > h1").textContent.trim(), "Happy Birthday!");
```

`p` 안에 `div` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("div > p"));
```

`p` 요소는 `class`인 `message`을 가져야 합니다.

```js
assert.exists(document.querySelector("div > p.message"));
```

`p` 요소는 `Wishing you all the happiness and joy on your special day!` 텍스트를 포함해야 합니다.

```js
assert.strictEqual(document.querySelector("div > p.message").textContent.trim(), "Wishing you all the happiness and joy on your special day!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css

```
