---
id: 6720b1f2680ec706ed534acf
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 다양한 가상 클래스와 가상 요소를 사용해 인사 카드 디자인을 연습합니다. HTML 보일러플레이트가 제공되어 있습니다.

워크숍을 시작하려면 `styles.css` 파일을 연결하세요.

# --hints--

`link` 요소는 `head` 요소 안에 있어야 합니다.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

`link` 요소는 `rel` 속성을 가지고 그 값은 `stylesheet` 여야 합니다.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

`link` 요소는 `href` 속성을 가지고 그 값은 `styles.css` 여야 합니다.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.match(linkHrefValue, /^(\.\/)?styles\.css$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
