---
id: 6720b7f3155a9278d6b6a8b8
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

이제 인사 카드에 스타일을 적용할 차례입니다.

`body` 요소에 대한 선택자를 추가한 다음:

- `font-family`을 `Arial`로 바꾸고 일반 `sans-serif`를 뒤에 붙이세요,
- 모든 면에 `40px` 패딩을 주세요,
- `text-align` 속성을 `center`로 설정하세요.

# --hints--

`body` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

`font-family` 선택자에서 `"Arial", sans-serif`을 `body`로 설정해야 합니다.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("font-family"), ["'Arial', sans-serif", "Arial, sans-serif", '"Arial", sans-serif']);
```

`padding` 선택자에서 `40px`을 `body`로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("padding"), "40px");
```

`text-align` 선택자에서 `center`를 `body`로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("text-align"), "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
  	</div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
