---
id: 6728fda085006772a247e779
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

`.card`에 모든 면에 `max-width`의 `400px`, `padding`의 `40px`, 그리고 위아래에는 `margin`의 `0`, 좌우에는 `auto`를 주되 (단축 속성을 사용하세요).

# --hints--

`.card`에는 `max-width`의 `400px`가 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("max-width"), "400px");
```

`.card`에는 `padding`의 `40px`가 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("padding"), "40px");
```

`.card`에는 `margin`의 `0 auto`가 있어야 합니다.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("margin"), ["0 auto", "0 auto 0 auto", "0 auto 0", "0px auto", "0px auto 0px auto", "0px auto 0px"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
  	</div>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
