---
id: 6729005c87b6320356ac1140
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

`.card` 요소에 스타일을 더 추가해야 합니다: `border-radius`의 `10px`과 `box-shadow` 값을 가진 `0 4px 8px gray`을 추가하세요.

# --hints--

`.card` 선택자에 `border-radius` 값을 가진 `10px` 속성이 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("border-radius"), "10px");
```

`.card` 선택자에 `box-shadow` 값을 가진 `0 4px 8px gray` 속성이 있어야 합니다.

```js
assert.strictEqual(
  new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("box-shadow"),
  "gray 0px 4px 8px" // this is the value outputted by getPropertyValue
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
  	</div>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
