---
id: 672b42241c0c6515f26786f4
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

이제 `div` 요소 아래에 새로운 `.message`를 추가하세요. 새 `div`는 `class` 속성이 `card-links`여야 합니다.

# --hints--

`#greeting-card` 안의 마지막 요소는 `div`이어야 합니다.

```js
assert.strictEqual(document.querySelector("#greeting-card").lastElementChild.tagName, "DIV");
```

새 `div`는 `class`가 `card-links`여야 합니다.

```js
assert.strictEqual(document.querySelector("#greeting-card").lastElementChild.className, "card-links");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      --fcc-editable-region--
      
      --fcc-editable-region--
  	</div>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
