---
id: 672b4434dd6d6a1edc405419
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

`a` 요소 안에 `.card-links` 요소 두 개를 추가하세요.

첫 번째 요소는 텍스트가 `Send Card`, `class`가 `send-link`, `href`가 `#send`여야 합니다.

두 번째 요소는 텍스트가 `Share on Social Media`, `class`가 `share-link`, `href`가 `#share`여야 합니다.

# --hints--

`a` 안의 첫 번째 `.card-links` 요소는 텍스트가 `Send Card`여야 합니다.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].innerText, "Send Card");
```

첫 번째 `a` 요소는 값이 `class`인 `send-link`를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].className, "send-link");
```

첫 번째 `a` 요소는 값이 `href`인 `#send`를 가져야 합니다.

```js
assert.match(document.querySelectorAll(".card-links > a")[0].href, /#send$/);
```

`a` 안의 두 번째 `.card-links` 요소는 텍스트가 `Share on Social Media`여야 합니다.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].innerText, "Share on Social Media");
```

두 번째 `a` 요소는 값이 `class`인 `share-link`를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].className, "share-link");
```

두 번째 `a` 요소는 값이 `href`인 `#share`를 가져야 합니다.

```js
assert.match(document.querySelectorAll(".card-links > a")[1].href, /#share$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
  	</div>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
