---
id: 67332fff82eb28128f08ec2c
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

`transform` 속성은 요소의 모양을 변형할 수 있습니다. 예를 들어, `scale(0.9)` 값을 주면 요소가 10% 작아집니다.

```css
p {
  transform: scale(0.9);
}
```

`transform` 선택자에 `.card:hover` 속성을 추가하고 `scale(1.1)`로 설정하세요.

# --hints--

`.card:hover` 선택자에는 `transform: scale(1.1)`가 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card:hover")?.getPropertyValue("transform"), "scale(1.1)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

.card:hover {
  background-color: khaki;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
