---
id: 6733352ce2a5642827b4f7f0
title: 18단계
challengeType: 0
dashedName: step-18
---

# --description--

`.card-links` 선택자를 추가하고 `margin-top` 속성을 `20px`로 설정하세요.

요소를 flexbox로 설정하려면 `display: flex`를 추가할 수 있습니다. flexbox에 대해서는 강의 후반부에 더 배우게 되며, 지금은 간단한 미리보기로 생각하세요.

두 링크 사이에 같은 간격이 있도록 하려면 `display` 속성을 `flex`로, 그리고 `justify-content`를 `space-around`로 설정하세요.

# --hints--

`.card-links` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".card-links"));
```

`.card-links` 선택자에는 `margin-top` 속성이 `20px`로 설정되어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links")?.getPropertyValue("margin-top"), "20px");
```

`.card-links` 선택자에는 `display` 속성이 `flex`로 설정되어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links")?.getPropertyValue("display"), "flex");
```

`.card-links` 선택자에는 `justify-content` 속성이 `space-around`로 설정되어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links")?.getPropertyValue("justify-content"), "space-around");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
  transition: transform 0.3s, background-color 0.3s ease
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

h1::before {
  content: "🥳 ";
}

h1::after {
  content: " 🥳";
}

.message {
  font-size: 1.2em;
  margin-bottom: 20px;
}

--fcc-editable-region--

--fcc-editable-region--
```
