---
id: 67332db0f793bf08f2a18528
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

두 번째 `section`를 채울 시간입니다! 

텍스트 `h2`를 포함하는 `#share` 요소에 `Sharing your card...` 요소를 추가한 다음, 텍스트 `p`가 있는 `Your card was shared on social media!` 요소를 추가하세요.

# --hints--

`#share` 요소는 `h2` 요소를 포함해야 합니다.

```js
assert.exists(document.querySelector("#share > h2"));
```

`h2` 요소는 텍스트 `Sharing your card...`를 포함해야 합니다.

```js
assert.strictEqual(document.querySelector("#share > h2").innerText.trim(), "Sharing your card...");
```

`#share` 내부의 두 번째 요소는 `p` 요소여야 합니다.

```js
assert.strictEqual(document.querySelector("#share").children[1].tagName, "P");
```

`p` 요소는 텍스트 `Your card was shared on social media!`를 가져야 합니다.

```js
assert.strictEqual(document.querySelector("#share").children[1].innerText.trim(), "Your card was shared on social media!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
  --fcc-editable-region--
      
  --fcc-editable-region--
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
