---
id: 67333082ab35e414eb4ea70d
title: 14단계
challengeType: 0
dashedName: step-14
---

# --description--

`a` 요소에 마우스를 올리면 배경색이 다른 색으로 전환됩니다. 이 전환이 어떻게 일어나는지 `transition` 속성으로 조절할 수 있습니다.

```css
a {
  transition: color 1s linear;
}
```

`transition` 속성이 받는 값은 순서대로 전환이 적용될 속성, 전환 지속 시간, 그리고 타이밍입니다.

전환이 적용되는 속성이 여러 개라면 각 값을 쉼표로 구분해 작성할 수 있습니다.

```css
p {
  transition: property1 0.1s, property2 0.6s linear;
}
```

만약 첫 번째 속성의 타이밍처럼 값이 생략되면 기본값이 적용됩니다.

`.card` 선택자에 `transition: transform 0.3s, background-color 0.3s ease`를 추가하세요.

직접 시도해 보세요, 호버 전환이 완성됩니다.

# --hints--

`.card` 선택자에 `transition` 속성을 `transform 0.3s, background-color 0.3s ease`로 설정해야 합니다.

```js
assert.oneOf(
  new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("transition"),
  [
    "transform 0.3s ease 0s, background-color 0.3s ease 0s", // this is what comes out of `getPropertyValue`
    "transform 0.3s, background-color 0.3s" // different environment got this instead
  ]
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
--fcc-editable-region--
  
--fcc-editable-region--
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

```
