---
id: 673332a6b63f9f1e1f81ba9b
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

이제 제목 뒤에 이모지도 추가하는 비슷한 작업을 할 수 있습니다.

`::after` 요소의 가상 요소 `h1`를 대상으로 하는 선택자를 만드세요. `content` 속성을 주고 값은 `" 🥳"`로 설정하세요(이모지 앞에 공백이 있다는 점을 참고하세요).

# --hints--

`h1::after` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("h1::after"));
```

`h1::after` 선택자에는 `content` 속성이 `" 🥳"`로 설정되어야 합니다.

```js
assert.strictEqual(
  new __helpers.CSSHelp(document).getStyle("h1::after")?.getPropertyValue("content"),
  `" 🥳"`
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
  transition: transform 0.3s, background-color 0.3s ease
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

h1::before {
  content: "🥳 ";
}

--fcc-editable-region--

--fcc-editable-region--
```
