---
id: 673337337f794d3025ded433
title: 19단계
challengeType: 0
dashedName: step-19
---

# --description--

`a` 내부의 `.card-links` 요소를 선택하고 다음을 지정하세요:

- `text-decoration` 속성을 `none`으로 설정하세요.
- `font-size` 속성을 `1em`로 설정하세요.
- `padding` 속성을 `10px 20px`로 설정하세요.
- `border-radius` 속성을 `5px`로 설정하세요.
- `color` 속성을 `white`로 설정하세요.
- `background-color` 속성을 `midnightblue`로 설정하세요.

# --hints--

`.card-links a` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".card-links a"));
```

`.card-links a` 선택자는 `text-decoration` 속성을 `none`로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links a")?.getPropertyValue("text-decoration"), "none");
```

`.card-links a` 선택자는 `font-size` 속성을 `1em`로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links a")?.getPropertyValue("font-size"), "1em");
```

`.card-links a` 선택자는 `padding` 속성을 `10px 20px`로 설정해야 합니다.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle(".card-links a")?.getPropertyValue("padding"), ["10px 20px", "10px 20px 10px", "10px 20px 10px 20px"]);
```

`.card-links a` 선택자는 `border-radius` 속성을 `5px`로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links a")?.getPropertyValue("border-radius"), "5px");
```

`.card-links a` 선택자는 `color` 속성을 `white`로 설정해야 합니다.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle(".card-links a")?.getPropertyValue("color").toLowerCase(), ["white", "#ffffff", "#fff", "rgb(255, 255, 255)", "hsl(0, 0%, 100%)"]);
```

`.card-links a` 선택자는 `background-color` 속성을 `midnightblue`로 설정해야 합니다.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle(".card-links a")?.getPropertyValue("background-color").toLowerCase(), ["midnightblue", "#191970", "rgb(25, 25, 112)", "hsl(240, 64%, 27%)"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
  transition: transform 0.3s, background-color 0.3s ease
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

h1::before {
  content: "🥳 ";
}

h1::after {
  content: " 🥳";
}

.message {
  font-size: 1.2em;
  margin-bottom: 20px;
}

.card-links {
  margin-top: 20px;
  display: flex;
  justify-content: space-around;
}

--fcc-editable-region--

--fcc-editable-region--
```
