---
id: 67333ba820f99c43904c9eff
title: 24단계
challengeType: 0
dashedName: step-24
---

# --description--

이미 방문한 `.card-links a` 요소를 대상으로 하는 새로운 선택자를 만드세요. 속성 `color`을 `crimson`로 설정하세요.

# --hints--

`.card-links a:visited` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".card-links a:visited"));
```

`.card-links a:visited` 선택자에 속성 `color`가 `crimson`로 설정되어야 합니다.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle(".card-links a:visited")?.getPropertyValue("color").toLowerCase(), ["crimson", "#dc143c", "rgb(220, 20, 60)", "hsl(348, 83%, 47%)"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
  transition: transform 0.3s, background-color 0.3s ease
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

h1::before {
  content: "🥳 ";
}

h1::after {
  content: " 🥳";
}

.message {
  font-size: 1.2em;
  margin-bottom: 20px;
}

.card-links {
  margin-top: 20px;
  display: flex;
  justify-content: space-around;
}

.card-links a {
  text-decoration: none;
  font-size: 1em;
  padding: 10px 20px;
  border-radius: 5px;
  color: white;
  background-color: midnightblue;
  transition: background-color 0.3s ease;
}

.card-links a:hover {
  background-color: orangered;
}

.card-links a:active {
  background-color: midnightblue;
}

.card-links a:focus {
  outline: 2px solid yellow;
}

--fcc-editable-region--

--fcc-editable-region--
```
