---
id: 68e4276483b0c71658aab1d0
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

Іноді вам потрібно буде створити новий рядок з іншим регістром.

Щоб створити новий рядок, у якому всі літери будуть у верхньому регістрі, ви можете використати метод `toUpperCase()`.

Ось приклад:

```js
const sentence = "I love to code!";
sentence.toUpperCase(); // "I LOVE TO CODE!"
```

Створіть змінну з назвою `upperCaseInput` і призначте їй результат виклику методу `toUpperCase()` на змінній `cleanedInput`.

# --hints--

У вас має бути змінна `upperCaseInput`.

```js
assert.exists(upperCaseInput);
```

Ваша змінна `upperCaseInput` має бути рядком.

```js
assert.isString(upperCaseInput);
```

Ви маєте призначити `cleanedInput.toUpperCase()` вашій змінній `upperCaseInput`.

```js
assert.equal(upperCaseInput, cleanedInput.toUpperCase());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

--fcc-editable-region--

--fcc-editable-region--
```
