---
id: 68e4276483b0c71658aab1cd
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Тепер настав час побачити результат.

Почніть із додавання `console.log()` з рядком `"Result of trimming whitespace from both ends:"`.

Потім нижче цього `console.log()` додайте ще один `console.log()` з змінною `cleanedInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Ви повинні вивести рядок `"Result of trimming whitespace from both ends:"` у консоль.

```js
assert.equal(getLogs()[2], "Result of trimming whitespace from both ends:");
```

Ви повинні вивести змінну `cleanedInput` у консоль.

```js
assert.equal(getLogs()[3], cleanedInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
--fcc-editable-region--

--fcc-editable-region--
```
