---
id: 68e4276483b0c71658aab1ce
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Іноді потрібно обрізати пробіли лише на початку рядка. Саме для цього зручно використовувати метод `trimStart()`.

Ось приклад:

```js
const sentence = "  I love learning.  ";
sentence.trimStart(); // "I love learning.  "
```

Пам’ятайте, що рядки є незмінними. Цей метод створить новий рядок із результатом, а не змінить оригінальний.

Створіть змінну `trimmedStart` і призначте їй результат застосування методу `trimStart()` до змінної `userInput`.

# --hints--

У вас має бути змінна `trimmedStart`.

```js
assert.exists(trimmedStart);
```

Ваша змінна `trimmedStart` має бути рядком.

```js
assert.isString(trimmedStart);
```

Ви маєте призначити `userInput.trimStart()` вашій змінній `trimmedStart`.

```js
assert.equal(trimmedStart, userInput.trimStart());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);
--fcc-editable-region--

--fcc-editable-region--
```
