---
id: 68e4276483b0c71658aab1cc
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Ваше значення `userInput` має зайві пробіли на обох кінцях рядка. Щоб створити новий рядок із видаленими цими пробілами, ви можете використати метод `trim()`.  

Ось приклад:

```js
const sentence = "  Learning about the trim method.  ";
sentence.trim(); // "Learning about the trim method."
```

Створіть змінну з ім’ям `cleanedInput` і призначте їй результат застосування методу `trim()` до змінної `userInput`.

# --hints--

У вас має бути змінна `cleanedInput`.

```js
assert.exists(cleanedInput);
```

Ваша змінна `cleanedInput` має бути рядком.

```js
assert.isString(cleanedInput);
```

Ви маєте призначити `userInput.trim()` вашій змінній `cleanedInput`.

```js
assert.equal(cleanedInput, userInput.trim());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);
--fcc-editable-region--

--fcc-editable-region--
```
