---
id: 68e4276483b0c71658aab1d2
title: Крок 13
challengeType: 1
dashedName: step-13
---

# --description--

Іноді вам доведеться використовувати кілька методів, щоб створити новий рядок. У наступних кількох кроках ви застосуєте деякі методи, які вже вивчили, щоб створити новий рядок, що використовує конвенцію іменування верблюжим регістром.

Почніть зі створення змінної з назвою `lowercaseWord` і призначте їй рядок `"camelcase"`.

Під цією змінною `lowercaseWord` створіть іншу змінну з назвою `camelCasedVersion` і поки що призначте їй порожній рядок.

Під цією змінною `camelCasedVersion` додайте `console.log()` з рядком `"Camel cased version:"`, а потім ще один `console.log()` для змінної `camelCasedVersion`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Ви повинні мати змінну з назвою `lowercaseWord`.

```js
assert.exists(lowercaseWord);
```

Ваш `lowercaseWord` має бути рядком.

```js
assert.isString(lowercaseWord);
```

Ви повинні призначити рядок `"camelcase"` змінній `lowercaseWord`.

```js
assert.equal(lowercaseWord, "camelcase");
```

Ви повинні мати змінну `camelCasedVersion`.

```js
assert.exists(camelCasedVersion);
```

Ваша змінна `camelCasedVersion` має бути рядком.

```js
assert.isString(camelCasedVersion);
```

Ви повинні призначити порожній рядок змінній `camelCasedVersion`.

```js
assert.isEmpty(camelCasedVersion);
```

Ви повинні вивести рядок `"Camel cased version:"` у консоль.

```js
assert.equal(getLogs()[12], "Camel cased version:");
```

Ви повинні вивести змінну `camelCasedVersion` у консоль.

```js
assert.equal(getLogs()[13], camelCasedVersion);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
