---
id: 68e4276483b0c71658aab1d1
title: Крок 12
challengeType: 1
dashedName: step-12
---

# --description--

Якщо потрібно створити новий рядок, у якому всі літери будуть у нижньому регістрі, ви можете використати метод `toLowerCase()`.  

Почніть зі створення змінної з назвою `lowerCaseInput` і призначте їй результат використання методу `toLowerCase()` на змінній `cleanedInput`. 

Потім виведіть у консоль рядок `"Result of using the toLowerCase() method:"`.

Нарешті, додайте ще один виклик `console.log()` для змінної `lowerCaseInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Ви повинні мати змінну `lowerCaseInput`.

```js
assert.exists(lowerCaseInput);
```

Ваша змінна `lowerCaseInput` має бути рядком.

```js
assert.isString(lowerCaseInput);
```

Ви повинні призначити `cleanedInput.toLowerCase()` вашій змінній `lowerCaseInput`.

```js
assert.equal(lowerCaseInput, cleanedInput.toLowerCase());
```

Ви повинні вивести рядок `"Result of using the toLowerCase() method:"` у консоль.

```js
assert.equal(getLogs()[10], "Result of using the toLowerCase() method:");
```

Ви повинні вивести змінну `lowerCaseInput` у консоль.

```js
assert.equal(getLogs()[11], lowerCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);


const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
