---
id: 68e4276483b0c71658aab1ca
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви попрактикуєтеся у роботі з різними методами рядків, які використовуються для форматування рядків.

Почніть зі створення змінної з назвою `userInput` і призначте їй таке значення: `"   Hello World!   "`.

# --hints--

У вас має бути змінна `userInput`.

```js
assert.exists(userInput);
```

Ваша змінна `userInput` має бути рядком.

```js
assert.isString(userInput);
```

Ви маєте призначити `"   Hello World!   "` вашій змінній `userInput`. Переконайтеся, що перед `H` і після `!` є по три пробіли.

```js
assert.equal(userInput, "   Hello World!   ");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
