---
id: 68e4276483b0c71658aab1d3
title: Крок 14
challengeType: 1
dashedName: step-14
---

# --description--

Як ви пам’ятаєте з попередніх уроків і практичних занять, ви можете витягувати частини рядка за допомогою методу `slice()`.

Ось нагадування:

```js
const str = "abcd";
str.slice(0, 3); // "abc"
```

Ви повинні змінити існуюче присвоєння `camelCasedVersion`, замінивши порожній рядок на `lowercaseWord.slice(0, 5)`.

Тепер у консолі має з’явитися слово `"camel"`.

# --hints--

Ваша змінна `camelCasedVersion` не повинна бути порожнім рядком.

```js
assert.isNotEmpty(camelCasedVersion);
```

Ви повинні призначити `lowercaseWord.slice(0, 5)` змінній `camelCasedVersion`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5));
```

Ви НЕ повинні призначати літеральний рядок `"camel"` змінній `camelCasedVersion`. Будь ласка, перечитайте інструкції для правильного підходу.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camel(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = ""
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
