---
id: 68f1179eaea3bb59ba434386
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Щоб обрізати пробіли в кінці рядка, ви можете використати метод `trimEnd()`.

Ось приклад:

```js
const sentence = "  I love learning.  ";
sentence.trimEnd(); // "  I love learning."
```

Створіть змінну з назвою `trimmedEnd` і призначте їй результат використання методу `trimEnd()` на змінній `userInput`.

# --hints--

У вас має бути змінна `trimmedEnd`.

```js
assert.exists(trimmedEnd);
```

Ваша змінна `trimmedEnd` має бути рядком.

```js
assert.isString(trimmedEnd);
```

Ви маєте призначити `userInput.trimEnd()` вашій змінній `trimmedEnd`.

```js
assert.equal(trimmedEnd, userInput.trimEnd());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

--fcc-editable-region--

--fcc-editable-region--
```
