---
id: 68f3f5ca1d2a490c97f14f03
title: Крок 9
challengeType: 1
dashedName: step-9
---

# --description--

Тепер настав час побачити результати обрізання пробілів у кінці рядка.

Під вашим змінним `trimmedEnd` додайте `console.log()` зі рядком `"After using the trimEnd() method, trailing spaces removed:"`.

Потім під цим `console.log()` додайте ще один `console.log()` зі змінною `trimmedEnd`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Ви маєте вивести рядок `"After using the trimEnd() method, trailing spaces removed:"` у консоль.

```js
assert.equal(getLogs()[6], "After using the trimEnd() method, trailing spaces removed:");
```

Ви маєте вивести змінну `trimmedEnd` у консоль.

```js
assert.equal(getLogs()[7], trimmedEnd);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
--fcc-editable-region--

--fcc-editable-region--
```
