---
id: 68f3f890010c870e17bbe754
title: Крок 11
challengeType: 1
dashedName: step-11
---

# --description--

Тепер настав час побачити результати.

Під вашим `upperCaseInput` змінним додайте `console.log()` з рядком `"Result of using the toUpperCase() method:"`.

Потім під цим `console.log()` додайте ще один `console.log()` зі змінною `upperCaseInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Ви повинні вивести рядок `"Result of using the toUpperCase() method:"` у консоль.

```js
assert.equal(getLogs()[8], "Result of using the toUpperCase() method:");
```

Ви повинні вивести змінну `upperCaseInput` у консоль.

```js
assert.equal(getLogs()[9], upperCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
--fcc-editable-region--

--fcc-editable-region--
```
