---
id: 68f3eeaa65614107ecd5f1f2
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Під вашим першим `console.log()` додайте ще один `console.log()`, щоб вивести змінну `userInput` у консоль.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Ви повинні вивести змінну `userInput` у консоль.

```js
assert.equal(getLogs()[1], userInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
--fcc-editable-region--

--fcc-editable-region--
```
