---
id: 68f3f3d65d51d80ac0b412fc
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Тепер настав час побачити результати обрізання пробілів на початку рядка.

Під вашим змінним `trimmedStart` додайте `console.log()` зі рядком `"After using the trimStart() method, leading spaces removed:"`.

Потім під цим `console.log()` додайте ще один `console.log()` зі змінною `trimmedStart`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Ви повинні вивести рядок `"After using the trimStart() method, leading spaces removed:"` у консоль.

```js
assert.equal(getLogs()[4], "After using the trimStart() method, leading spaces removed:");
```

Ви повинні вивести змінну `trimmedStart` у консоль.

```js
assert.equal(getLogs()[5], trimmedStart);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
--fcc-editable-region--

--fcc-editable-region--
```
