---
id: 68f405c8c9df44132c8fe1bb
title: Крок 15
challengeType: 1
dashedName: step-15
---

# --description--

Як ви пам’ятаєте з попередніх уроків, конвенція іменування верблюжим регістром починається з літери в нижньому регістрі для першого слова, а кожне наступне слово починається з великої літери, за якою йдуть літери в нижньому регістрі.

Ось кілька нагадувальних прикладів:

```md
freeCodeCamp
camelCase
learningIsFun
```

Друге слово у змінній `lowercaseWord` — це `"case"`. Щоб отримати доступ до `c` у цьому слові, ви можете використати `lowercaseWord[5]`.

Оновіть присвоєння `camelCasedVersion`, щоб використати оператор `+` для конкатенації `lowercaseWord.slice(0, 5)` з результатом застосування правильного методу для перетворення рядків у верхній регістр до `lowercaseWord[5]`.

Тепер у консолі ви повинні побачити `camelC`.

# --hints--

Ви повинні конкатенувати `lowercaseWord.slice(0, 5)` з `lowercaseWord[5].toUpperCase()`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase());
```

Ви НЕ повинні призначати літеральний рядок `"camelC"` змінній `camelCasedVersion`. Будь ласка, перечитайте інструкції для правильного підходу.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelC(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5)
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
