---
id: 68f40925d810ae14b2d5a169
title: Крок 16
challengeType: 1
dashedName: step-16
---

# --description--

Зараз ваша змінна `camelCasedVersion` містить значення `camelC`. Але бажаний результат — `camelCase`.

Щоб отримати решту букв `ase`, ви можете використати метод `slice()` з від’ємним індексом. Нагадаємо, що використання від’ємних індексів каже методу `slice()` витягувати символи з кінця.

Ось нагадування у вигляді прикладу:

```js
const str = "abcd";
str.slice(-2); // cd
```

Оновіть присвоєння `camelCasedVersion`, щоб також конкатенувати `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` з результатом використання методу `slice()` на `lowercaseWord`. Передайте числове значення `-3` для методу `slice()`.

І з цією останньою зміною ваше практичне заняття завершене!

# --hints--

Ви повинні конкатенувати `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` з `lowercaseWord.slice(-3)`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase() + lowercaseWord.slice(-3));
```

Ви НЕ повинні присвоювати літеральний рядок `"camelCase"` змінній `camelCasedVersion`. Будь ласка, перечитайте інструкції для правильного підходу.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelCase(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

# --solutions--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";

const camelCasedVersion =
  lowercaseWord.slice(0, 5) +
  lowercaseWord[5].toUpperCase() +
  lowercaseWord.slice(-3);

console.log("Camel cased version:");
console.log(camelCasedVersion);
```
