---
id: 641da7bfbc7f0f477438ad8a
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

이제 `displayAuthors` 함수를 호출할 차례입니다. 하지만 다시 한 번, 모든 저자를 한꺼번에 페이지에 표시하고 싶지 않습니다. 대신 `startingIndex`과 `endingIndex` 변수를 사용해 저자 중 일부만 추출할 수 있습니다. 이를 가장 잘 수행하는 방법은 `.slice()` 배열 메서드입니다.

먼저 `authorDataArr`를 출력하는 콘솔 로그 문을 제거하세요. 그런 다음 `displayAuthors` 배열과 `authorDataArr`을 인수로 하여 `.slice()` 함수를 호출하세요. 시작점에는 `startingIndex` 변수를, 끝점에는 `endingIndex` 변수를 사용하세요.

# --before-all--

```js
window.fetch = () => Promise.resolve({json: () => Promise.resolve([{ author: 'Whoever', image: 'data:image/gif;base64,R0lGODlhAQABAIAAAAAAAP///ywAAAAAAQABAAACAUwAOw==', url: "http://not-a-real-url.nowhere/", bio: 'words go here' }])});
```

# --hints--

`data` 변수에 `authorDataArr`을 할당해야 합니다

```js
assert.match(code, /authorDataArr\s*=\s*data\s*;?/)
```

`authorDataArr`를 출력하는 콘솔 로그 문을 제거해야 합니다

```js
assert.notMatch(code, /console\.log\(\s*("|'|`)Author\s+Data\s+Array:\s*\1\s*,\s*authorDataArr\s*\)\s*;?/)
```

`displayAuthors` 함수를 호출해야 합니다.

```js
assert.match(code, /displayAuthors\(/)
```

`displayAuthors` 함수를 호출하고 `authorDataArr.slice()` 인수를 전달해야 합니다.

```js
assert.match(code, /displayAuthors\(\s*authorDataArr\.slice\(/)
```

`slice()` 메서드의 첫 번째 매개변수는 `startingIndex`이어야 합니다.

```js
assert.match(code, /displayAuthors\(\s*authorDataArr\.slice\(\s*startingIndex/)
```

`slice()` 메서드의 두 번째 매개변수는 `endingIndex`이어야 합니다.

```js
assert.match(code, /displayAuthors\(\s*authorDataArr\.slice\(\s*startingIndex\s*,\s*endingIndex\s*\)\s*\)\s*;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

let startingIndex = 0;
let endingIndex = 8;
let authorDataArr = [];

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    --fcc-editable-region--
    authorDataArr = data;
    console.log("Author Data Array:", authorDataArr);

    --fcc-editable-region-- 
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });

const displayAuthors = (authors) => {
  authors.forEach(({ author, image, url, bio }, index) => {
    authorContainer.innerHTML += `
    <div id="${index}" class="user-card">
      <h2 class="author-name">${author}</h2>
    </div>
  `;
  });
};
```
