---
id: 641da5462576784453146ec2
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

원하는 데이터를 얻었으니, 이를 사용해 UI를 채울 수 있습니다. 하지만 가져온 데이터에는 26명의 저자가 배열로 들어 있고, 모두 한꺼번에 페이지에 추가하면 성능이 저하될 수 있습니다.

대신 한 번에 8명의 저자만 추가하고, 더 이상 표시할 데이터가 없을 때까지 8명씩 더 추가하는 버튼을 만들어야 합니다.

`let`를 사용해 `startingIndex`과 `endingIndex`라는 이름의 변수를 만들고, 각각 숫자 값 `0`과 `8`를 할당하세요. 또한 `authorDataArr`라는 변수를 `let`로 선언하고 빈 배열로 설정하세요.

# --before-all--

```js
window.fetch = () => Promise.resolve({json: () => Promise.resolve([{ author: 'Whoever', image: 'data:image/gif;base64,R0lGODlhAQABAIAAAAAAAP///ywAAAAAAQABAAACAUwAOw==', url: "http://not-a-real-url.nowhere/", bio: 'words go here' }])});
```

# --hints--

`let`를 사용해 `startingIndex`라는 이름의 변수를 선언하세요.

```js
assert.match(code, /let\s+startingIndex/)
```

`startingIndex` 변수를 `0`으로 설정하세요.

```js
assert.match(code, /let\s+startingIndex\s*=\s*0\s*;?/)
```

`let`를 사용해 `endingIndex`라는 이름의 변수를 선언하세요.

```js
assert.match(code, /let\s+endingIndex/)
```

`endingIndex` 변수를 `8`로 설정하세요.

```js
assert.match(code, /let\s+endingIndex\s*=\s*8\s*;?/)
```

`let`를 사용해 `authorDataArr`라는 이름의 변수를 선언하세요.

```js
assert.match(code, /let\s+authorDataArr/)
```

`authorDataArr` 변수를 빈 배열(`[]`)로 설정하세요.

```js
assert.match(code, /let\s+authorDataArr\s*=\s*\[\s*\]\s*;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

--fcc-editable-region--

--fcc-editable-region--

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    console.log(data);   
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });
```
