---
id: 641da86294fd9f485d3c2bf0
title: 19단계
challengeType: 0
dashedName: step-19
---

# --description--

다음으로, freeCodeCamp 뉴스의 작성자 페이지에 링크를 추가하세요.

`class` `"author-link"`를 가진 앵커 요소를 추가하고, `url` 속성의 값으로 `href`를 보간하며, `target`를 `"_blank"`로 설정하세요. 앵커 요소의 텍스트로는 `author`를 보간한 다음 `"'s author page"` 텍스트를 넣으세요. 예를 들어, `"Quincy Larson's author page"`입니다.

# --before-all--

```js
window.fetch = () => Promise.resolve({json: () => Promise.resolve([{ author: 'Whoever', image: 'data:image/gif;base64,R0lGODlhAQABAIAAAAAAAP///ywAAAAAAQABAAACAUwAOw==', url: "http://not-a-real-url.nowhere/", bio: 'words go here' }])});
```

# --hints--

앵커 요소를 만들어야 합니다.

```js
await __helpers.retryingTest(() => document.querySelector('a'), "'a' element not found");
```

앵커 요소에 `"author-link"` 클래스를 지정해야 합니다.

```js
assert.include(document.querySelector('a')?.className, "author-link");
```

앵커의 `url` 속성 값으로 `href`를 보간해야 합니다.

```js
assert.equal(document.querySelector('a')?.getAttribute('href'), authorDataArr[0].url);
```

앵커 요소의 `target` 속성을 `"_blank"`로 설정해야 합니다.

```js
assert.equal(document.querySelector('a')?.getAttribute('target'), '_blank');
```

앵커 요소 안에 `author`를 보간한 다음 `"'s author page"` 텍스트를 넣어야 합니다.

```js
assert.equal(document.querySelector('a')?.textContent.trim(), `${authorDataArr[0].author}'s author page`);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

let startingIndex = 0;
let endingIndex = 8;
let authorDataArr = [];

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    authorDataArr = data;
    displayAuthors(authorDataArr.slice(startingIndex, endingIndex));   
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });

const displayAuthors = (authors) => {
  authors.forEach( ({ author, image, url, bio }, index) =>
  (authorContainer.innerHTML += `
    <div id="${index}" class="user-card">
      <h2 class="author-name">${author}</h2>
      <img class="user-img" src="${image}" alt="${author} avatar" />
      <p class="bio">${bio}</p>
      --fcc-editable-region--

      --fcc-editable-region--
    </div>
  `)
  )
};
```
