---
id: acda2fb1324d9b0fa741e6b5
title: بناء أداة لتأكيد نهاية السلسلة النصية
challengeType: 26
dashedName: implement-a-string-ending-checker-function
---

# --description--

في هذا المختبر، ستنفذ دالة تتحقق مما إذا كانت سلسلة نصية تنتهي بالسلسلة النصية الهدف المعطاة.

**الهدف:** حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ دالة باسم `confirmEnding` تأخذ معلمتين: السلسلة النصية التي ستتحقق منها والسلسلة النصية التي ستتحقق مقابلها.
2. يجب أن تُرجع الدالة `true` إذا انتهت السلسلة النصية الأولى بالسلسلة النصية الثانية، و`false` خلاف ذلك.
3. لا يجب أن تستخدم الطريقة `.endsWith()`؛ بدلاً من ذلك، استخدم إحدى طرق السلاسل النصية الفرعية في JavaScript لتحقيق ذلك.

# --hints--

يجب أن تنشئ دالة باسم `confirmEnding`.

```js
assert.isFunction(confirmEnding);
```

يجب أن تأخذ `confirmEnding` معلمتين.

```js
assert.lengthOf(confirmEnding, 2);
```

يجب أن تُرجع `confirmEnding("Bastian", "n")` القيمة `true`.

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

يجب أن تُرجع `confirmEnding("Congratulation", "on")` القيمة `true`.

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

يجب أن تُرجع `confirmEnding("Connor", "n")` القيمة `false`.

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

يجب أن تُرجع `confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")` القيمة `false`.

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

يجب أن تُرجع `confirmEnding("He has to give me a new name", "name")` القيمة `true`.

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

يجب أن تُرجع `confirmEnding("Open sesame", "same")` القيمة `true`.

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

يجب أن تُرجع `confirmEnding("Open sesame", "sage")` القيمة `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

يجب أن تُرجع `confirmEnding("Open sesame", "game")` القيمة `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

يجب أن تُرجع `confirmEnding("If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing", "mountain")` القيمة `false`.

```js
assert.isFalse(
  confirmEnding(
    "If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing",
    'mountain'
  )
);
```

يجب أن تُرجع `confirmEnding("Abstraction", "action")` القيمة `true`.

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

يجب ألا تستخدم شفرتك الطريقة المدمجة `.endsWith()` لحل المختبر.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
