---
id: 68e4c81577243e4db1676491
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、意味論的なHTMLを使って、Quincy Larsonの開発者の仕事を得るためのヒントを含むウェブページを作成する練習をします。基本的なHTMLのボイラープレートは用意されています。

まず、テキストが`h1`の`Quincy's Tips for Getting a Developer Job`要素を作成してください。

# --hints--

`h1`要素があるはずです。

```js
assert.exists(document.querySelector('h1'));
```

`h1`要素にはテキスト`Quincy's Tips for Getting a Developer Job`が含まれているはずです。

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), "Quincy's Tips for Getting a Developer Job")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
