---
id: 68e65ec5348e40a739903e6d
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`h1` 要素の下に段落要素を作成してください。この段落の中にテキスト `Learning to code is hard, but as Quincy Larson says, You can become a developer.` を追加してください。

# --hints--

`p` 要素があるはずです。

```js
assert.exists(document.querySelector('p'));
```

段落にはテキスト `Learning to code is hard, but as Quincy Larson says, You can become a developer.` が含まれている必要があります。スペルをよく確認してください。

```js
assert.equal(document.querySelector('p')?.textContent.trim(), "Learning to code is hard, but as Quincy Larson says, You can become a developer.");
```

段落は `h1` 要素の下にある必要があります。

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
