---
id: 68e964b84f954c322376e592
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

引用元がウェブサイトの場合、`cite` 属性で引用元を示すことができます。この属性の値は有効なURLである必要があります。この属性は引用ブロックの表示を変えませんが、スクリーンリーダーや検索エンジンに引用についての情報を提供するのに非常に役立ちます。

以下は`cite` 属性を持つインライン引用要素の例です。

```html
<p>
  Nancy said,
  <q cite="https://example.com">Learning is fun!</q>
</p>
```

このURLを使ってインライン引用要素に`cite` 属性を追加してください。

`https://www.freecodecamp.org/news/learn-to-code-book/`

# --hints--

あなたの`q` 要素には`cite` 属性が必要です。

```js
const qEl = document.querySelector('h1 + p > q');
assert.exists(qEl?.getAttribute('cite'));
```

あなたの`cite` 要素の`q` 属性は`https://www.freecodecamp.org/news/learn-to-code-book/` の値であるべきです。

```js
const qEl = document.querySelector('h1 + p > q');
assert.equal(qEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
--fcc-editable-region--
      <q>You can become a developer.</q>
--fcc-editable-region--
    </p>
  </body>
</html>
```
