---
id: 68e90335bd53a901c41dfc13
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

前のレッスンで、`q` 要素が引用されたテキストを周囲の内容と区別するために使われることを学びました。

例を示します：

```html
<p>
  Nancy said, <q>Learning is fun!</q>
</p>
```

ほとんどの最新ブラウザーは、`q` 要素を使うとインライン引用の周りに自動的に引用符を追加します。

段落内で、段落の他の部分は変更せずに `You can become a developer.` をインライン引用要素で囲んでください。

# --hints--

`q` 要素があるはずです。

```js
assert.exists(document.querySelector('q'));
```

`q` 要素は、`h1` 要素の下にある段落の中にネストされている必要があります。

```js
assert.exists(document.querySelector('h1 + p > q'));
```

`q` 要素にはテキスト `You can become a developer.` が含まれている必要があります。

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent.trim(), 'You can become a developer.');
```

`q` 要素の中のテキストは、先頭や末尾に空白（スペース、タブ、改行）があってはいけません。

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent, 'You can become a developer.');
```

元のテキストを変更せずに、文の適切な部分の周りにのみ `q` タグを追加してください。`p` 要素はまだ `Learning to code is hard, but as Quincy Larson says, You can become a developer.` と表示されているはずです。

```js
const pElCleanedText = document.querySelector('h1 + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, 'Learning to code is hard, but as Quincy Larson says, You can become a developer.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    <p>Learning to code is hard, but as Quincy Larson says, You can become a developer.</p>
--fcc-editable-region--
  </body>
</html>
```
