---
id: 68ea7e7b916e8270483689b6
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

`h2` 要素の下に、クインシーの別の引用を追加します。今回は、その引用は大きな段落の一部ではなく、段落全体が引用になります。このような引用文を区別するために、引用ブロック要素 `blockquote` を使います。ブラウザでは、テキストが少しインデントされて表示されます。

以下は、引用テキストを含む引用ブロック要素の例です：

```html
<blockquote>
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

では、最初のセクション内で、`h2` 要素の下に、テキスト `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?` を含む引用ブロック要素を追加してください。

# --hints--

最初の `section` 要素の中に、`blockquote` 要素の下に `h2` 要素があるはずです。

```js
assert.exists(document.querySelector('section:first-of-type > h2 + blockquote'));
```

`blockquote` 要素にはテキスト `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?` が含まれているはずです。

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.innerText.trim(), 'Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
