---
id: 68ea5ccc1229b85eca38a3e8
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

最初の`section`要素の中に、テキスト`h2`を持つ`Envisioning Success`要素を追加してください。

# --hints--

最初の`section`要素の中に、`h2`要素があるはずです。

```js
assert.exists(document.querySelector('section:first-of-type > h2'));
```

あなたの`h2`はテキスト`Envisioning Success`を持つべきです。

```js
const firstH2El = document.querySelector('section:first-of-type > h2');
assert.equal(firstH2El?.innerText.trim(), `Envisioning Success`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
