---
id: 68eaa9ef8a0867071f195f59
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

これまで`cite`属性を使って引用元を示してきましたが、その属性はユーザーに引用元を視覚的に示すものではありません。

引用元を視覚的に示したい場合は、引用ブロック要素の外側に`cite`という引用要素を追加できます。引用要素は、本、記事、曲、映画、ウェブサイト、研究論文などの参照された創作物のタイトルをマークアップするためのHTML要素です。

例を示します。

```html
<div>
  <blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
    The first thing you should consider about education is this is an economic decision.
  </blockquote>
  <p>&mdash;Quincy Larson, <cite>Is College Still Worth it? Tips from my 20 Years in Adult Education</cite></p>
</div>
```

引用ブロック要素の下にある`p`要素の中で、`How to Learn to Code and Get a Developer Job [Full Book]`を`cite`要素で囲んでください。

# --hints--

`cite`要素の下にある`p`要素の中に`blockquote`要素があるはずです。

```js
assert.exists(document.querySelector('blockquote + p > cite'));
```

`cite`要素には`How to Learn to Code and Get a Developer Job [Full Book]`というテキストが含まれているはずです。

```js
const citeEl = document.querySelector('blockquote + p > cite');
assert.equal(citeEl?.innerText.trim(), 'How to Learn to Code and Get a Developer Job [Full Book]');
```

追加したばかりの`cite`要素を無視すると、`p`要素のテキストはまだ`—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]`のままです。

```js
const pElCleanedText = document.querySelector('blockquote + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, '—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
      --fcc-editable-region--
          &mdash;Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]
      --fcc-editable-region--
        </p>
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
