---
id: 68ebe9cd7523d340c0b343d3
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

3つ目で最後の`section`要素の中に、前のセクションで作成したのと同じHTML構造を作成します。

まず、最後のセクションの中に、テキストが`h2`の`Importance of Building a Reputation`要素を追加してください。

次に、`h2`要素の下に、`cite`属性が`https://www.freecodecamp.org/news/learn-to-code-book/`に設定されたブロック引用要素を追加します。

最後に、そのブロック引用要素の中に、以下のテキストを持つ3つの段落要素を順番にネストしてください。

1. `Share short video demos of your projects.`
2. `Keep applying to speak at bigger and bigger conferences.`
3. `Hang out at hackerspaces and help people who are even newer to coding than you.`

おめでとうございます！これでこのワークショップは完了です。

# --hints--

3つ目のセクションの中に`h2`要素がネストされているはずです。

```js
assert.exists(document.querySelector('main > section:nth-of-type(3) > h2'));
```

`h2`要素のテキストは`Importance of Building a Reputation`であるべきです。

```js
const h2El = document.querySelector('main > section:nth-of-type(3) > h2');
assert.equal(h2El?.innerText.trim(), 'Importance of Building a Reputation');
```

3つ目のセクションの中に`blockquote`要素がネストされているはずです。

```js
assert.exists(document.querySelector('main > section:nth-of-type(3) > blockquote'));
```

`blockquote`要素は`h2`要素の下にあるべきです。

```js
assert.exists(document.querySelector('main > section:nth-of-type(3) > h2 + blockquote'));
```

`blockquote`要素は`cite`属性を持つべきです。

```js
const blockquoteEl = document.querySelector('main > section:nth-of-type(3) > blockquote');
assert.exists(blockquoteEl?.getAttribute('cite'));
```

`cite`要素の`blockquote`属性は`https://www.freecodecamp.org/news/learn-to-code-book/`の値を持つべきです。

```js
const blockquoteEl = document.querySelector('main > section:nth-of-type(3) > blockquote');
assert.equal(blockquoteEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

3つ目のセクションの中の`p`要素に3つの`blockquote`要素がネストされているはずです。

```js
const pElements = document.querySelectorAll('main > section:nth-of-type(3) > blockquote > p');
assert.lengthOf(pElements, 3);
```

最初の`p`要素はテキスト`Share short video demos of your projects.`を持つべきです。

```js
const firstPEl = document.querySelector('main > section:nth-of-type(3) > blockquote > p:nth-of-type(1)');
assert.equal(firstPEl?.innerText.trim(), 'Share short video demos of your projects.');
```

2番目の`p`要素はテキスト`Keep applying to speak at bigger and bigger conferences.`を持つべきです。

```js
const secondPEl = document.querySelector('main > section:nth-of-type(3) > blockquote > p:nth-of-type(2)');
assert.equal(secondPEl?.innerText.trim(), "Keep applying to speak at bigger and bigger conferences.");
```

3番目の`p`要素はテキスト`Hang out at hackerspaces and help people who are even newer to coding than you.`を持つべきです。

```js
const thirdPEl = document.querySelector('main > section:nth-of-type(3) > blockquote > p:nth-of-type(3)');
assert.equal(thirdPEl?.innerText.trim(), 'Hang out at hackerspaces and help people who are even newer to coding than you.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
          &mdash;Quincy Larson, <cite>How to Learn to Code and Get a Developer Job [Full Book]</cite>
        </p>
      </section>
      <section>
        <h2>Importance of Networking</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          <p>So much of getting a job is who you know.</p>
          <p>It's OK to be an introvert, but you do need to push your boundaries.</p>
          <p>Create GitHub, Twitter, LinkedIn, and Discord accounts.</p>
          <p>Go to tech meetups and conferences. Travel if you have to.</p>
        </blockquote>
      </section>
      <section>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
          &mdash;Quincy Larson, <cite>How to Learn to Code and Get a Developer Job [Full Book]</cite>
        </p>
      </section>

      <section>
        <h2>Importance of Networking</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          <p>So much of getting a job is who you know.</p>
          <p>It's OK to be an introvert, but you do need to push your boundaries.</p>
          <p>Create GitHub, Twitter, LinkedIn, and Discord accounts.</p>
          <p>Go to tech meetups and conferences. Travel if you have to.</p>
        </blockquote>
      </section>

      <section>
        <h2>Importance of Building a Reputation</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          <p>Share short video demos of your projects.</p>
          <p>Keep applying to speak at bigger and bigger conferences.</p>
          <p>Hang out at hackerspaces and help people who are even newer to coding than you.</p>
        </blockquote>
      </section>
    </main>
  </body>
</html>
```
