---
id: 66c051d13a6a20255a963695
title: Eine stilisierte To-Do-Liste erstellen
challengeType: 25
dashedName: build-a-stylized-to-do-list
demoType: onClick
---

# --description--

In diesem Labor üben Sie die verschiedenen Stile, die auf Links angewendet werden können, wenn sie mit der Maus darübergefahren, fokussiert, angeklickt oder besucht werden.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine ungeordnete Liste mit der Klasse `todo-list` haben.
2. Innerhalb der ungeordneten Liste sollten vier Listenelemente sein.
3. Innerhalb jedes Listenelements sollte Folgendes enthalten sein:

   - Ein `input`-Element mit dem Typ `checkbox` und `id`, das auf einen eindeutigen Wert gesetzt ist.
   - Ein `label`-Element mit dem `for`-Attribut, das auf das entsprechende `input`-Element-`id` verweist.
   - Eine ungeordnete Liste mit der Klasse `sub-item`.
   - Ein Listenelement mit einem Ankerelement darin. Der Anker sollte die Klasse `sub-item-link`, einen gültigen `href`-Wert und einen `target`-Wert haben, der den Link in einem neuen Tab öffnet.

4. Ihre `a`-Elemente sollten keine Textdekorationen haben.
5. Sie sollten die Textfarbe von nicht besuchten Links auf eine Farbe Ihrer Wahl setzen.
6. Wenn Ihre Links besucht sind, sollte sich die Farbe in eine andere Farbe Ihrer Wahl ändern.
7. Wenn Sie mit der Maus über Ihre Links fahren, sollte sich die Farbe in eine andere Farbe Ihrer Wahl ändern.
8. Wenn Ihre Links den Fokus haben, sollte eine farbige Umrandung um den Link erscheinen.
9. Wenn Ihre Links angeklickt werden, sollte sich die Farbe in eine andere Farbe Ihrer Wahl ändern.

**Hinweis:** Vergewissern Sie sich, dass Sie Ihr Stylesheet in Ihrem HTML verlinken und Ihr CSS anwenden.

# --hints--

Sie sollten eine ungeordnete Liste mit der Klasse `todo-list` haben.

```js
assert.exists(document.querySelector('ul.todo-list'));
assert.lengthOf(document.querySelectorAll('ul.todo-list'), 1);
```

Sie sollten vier Listenelemente innerhalb der ungeordneten Liste haben.

```js
const li = document.querySelectorAll('ul.todo-list > li');
assert.lengthOf(li, 4);
```

Das `li` innerhalb des `ul` mit der Klasse `todo-list` sollte ein `input`-Element mit dem `type` von `checkbox` enthalten.

```js
const checkboxes = document.querySelectorAll(
  'ul.todo-list li input[type="checkbox"]'
);
assert.lengthOf(checkboxes, 4);
```

Das `li` innerhalb des `ul` mit der Klasse `todo-list` sollte ein `label`-Element enthalten.

```js
const labels = document.querySelectorAll(
  'ul.todo-list li label'
);
assert.lengthOf(labels, 4);
```

Alle `input`-Elemente sollten ein `id` haben.

```js
const inputs = document.querySelectorAll('input');
assert.isAbove(inputs.length, 0);
for (let input of inputs) {
  assert.isAbove(input?.id.length, 0);
}
```

Alle `label`-Elemente sollten ein `for`-Attribut haben.

```js
const labels = document.querySelectorAll('label');
assert.isAbove(labels.length, 0);
for (let label of labels) {
  assert.isAbove(label.htmlFor.length, 0);
}
```

Alle `label`-Elemente sollten Text enthalten.

```js
const labels = document.querySelectorAll('label');
assert.isAbove(labels.length, 0);
for (let label of labels) {
  assert.isAbove(label.innerText.length, 0);
}
```

Die `id`- und `for`-Attribute der `input`- und `label`-Elementpaare sollten entsprechende Werte haben.

```js
const labels = document.querySelectorAll("label");
assert.isNotEmpty(labels);

document.querySelectorAll("input").forEach((input, index) => {
  assert.equal(input.id, labels[index].htmlFor);
});
```

Nach den `label`-Elementen sollte eine ungeordnete Liste mit der Klasse `sub-item` folgen.

```js
const labelSbItems = document.querySelectorAll('ul.todo-list li label + ul.sub-item');
assert.lengthOf(labelSbItems, 4);
```

Das `li` innerhalb des `ul` mit der Klasse `sub-item` sollte ein Ankerelement mit der Klasse `sub-item-link` haben.

```js
const allItems = document.querySelectorAll(
  'ul.todo-list li label + ul.sub-item li a.sub-item-link'
);
assert.lengthOf(allItems, 4);
```

Alle Ankerelemente sollten ein gültiges `href`-Attribut haben.

```js
const links = document.querySelectorAll('a');
assert.isAbove(links.length, 0);
for (let link of links) {
  assert.isAbove(link.href.length, 0);
}
```

Alle Ankerelemente sollten Text enthalten.

```js
const links = document.querySelectorAll('a');
assert.isAbove(links.length, 0);
for (let link of links) {
  assert.isAbove(link.innerText.length, 0);
}
```

Jedes `a`-Element sollte ein `target`-Attribut mit dem Wert `_blank` haben.

```js
const anchors = document.querySelectorAll('a');
assert.isAbove(anchors.length, 0);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

Ihre `a`-Elemente sollten keine Textdekorationen haben.

```js
assert.isTrue(
  new __helpers.CSSHelp(document).getStyle('a.sub-item-link')
    ?.textDecoration === 'none' ||
    new __helpers.CSSHelp(document).getStyle('a')?.textDecoration === 'none' ||
    new __helpers.CSSHelp(document).getStyle('.sub-item-link')
      ?.textDecoration === 'none'
);
```

Nicht besuchte Links sollten eine Textfarbe haben, die Sie wählen.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:link')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:link')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:link')?.color
);
```

Die Links sollten ihre Farbe ändern, wenn man mit der Maus darüberfährt.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:hover')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:hover')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:hover')?.color
);
```

Die Links sollten ihre Farbe ändern, wenn sie angeklickt werden.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:active')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:active')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:active')?.color
);
```

Die Links sollten eine Umrandung haben, wenn sie den Fokus haben.

```js
const css = new __helpers.CSSHelp(document);
assert.exists(
  css.getStyle('.sub-item-link:focus')?.outline ||
    css.getStyle('a:focus')?.outline ||
    css.getStyle('a.sub-item-link:focus')?.outline ||
    css.getStyle('.sub-item-link:focus')?.border ||
    css.getStyle('a:focus')?.border ||
    css.getStyle('a.sub-item-link:focus')?.border
);
```

Die Links sollten ihre Farbe ändern, wenn sie besucht wurden.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:visited')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:visited')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:visited')?.color
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Styled To-Do List</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Styled To-Do List</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div class="container">
      <h1>My To-Do List</h1>
      <ul class="todo-list">
        <li>
          <input type="checkbox" id="task1" />
          <label for="task1">Explore gaming keyboards</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link">Store link</a>
            </li>
          </ul>
        </li>

        <li>
          <input type="checkbox" id="task2" />
          <label for="task2">Finish the report</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link"
                >Request Access</a
              >
            </li>
          </ul>
        </li>

        <li>
          <input type="checkbox" id="task3" />
          <label for="task3">View Phone's Warranty</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link"
                >View Receipts</a
              >
            </li>
          </ul>
        </li>

        <li>
          <input type="checkbox" id="task4" />
          <label for="task4">Check Processor Specs</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link"
                >View Model Number</a
              >
            </li>
          </ul>
        </li>
      </ul>
    </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: #f0f0f0;
  margin: 0;
  padding: 20px;
}

.container {
  max-width: 400px;
  margin: 0 auto;
  background-color: #fff;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
  color: #333;
}

.todo-list {
  list-style-type: none;
  padding: 0;
}

.todo-list li {
  margin-bottom: 15px;
  line-height: 1.6;
  padding: 10px;
  background-color: #e5f1e5;
  border-radius: 10px;
  position: relative;
}

a.sub-item-link {
  margin-left: 10px;
  text-decoration: none;
}

.sub-item-link:link {
  color: blue;
}

.sub-item-link:visited {
  color: purple;
}

.sub-item-link:hover {
  color: red;
}

.sub-item-link:focus {
  outline: 2px solid blue;
  outline-offset: 2px;
}

.sub-item-link:active {
  color: green;
}

.sub-item {
  list-style-type: square;
}
```
