---
id: 6a57c331ae94fa4631b60862
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Кожен артефакт також відстежує, де він був, як його описують і чи він виставлений на показ.

Додайте ще три пропи до артефакту, що зберігається під `101`:

- `locations`: масив, що містить два об’єкти, `{ gallery: "Hall A", year: 2020 }` та `{ gallery: "Hall C", year: 2024 }`
- `tags`: масив, що містить рядки `"gold"` і `"egypt"`
- `onDisplay`: булевий `true`

# --hints--

`locations` артефакту має бути масивом із двох заданих об’єктів локацій.

```js
assert.deepEqual(collection[101].locations, [
  { gallery: "Hall A", year: 2020 },
  { gallery: "Hall C", year: 2024 },
]);
```

`tags` артефакту має бути `["gold", "egypt"]`.

```js
assert.deepEqual(collection[101].tags, ["gold", "egypt"]);
```

`onDisplay` артефакту має бути `true`.

```js
assert.strictEqual(collection[101].onDisplay, true);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
