---
id: 6a57c331ae94fa4631b60868
title: Крок 9
challengeType: 1
dashedName: step-9
---

# --description--

Куратори позначають артефакти тегами, щоб їх було легше шукати, і тег ніколи не повинен додаватися двічі. Перед тим як додати новий тег, ви можете використати метод `includes()`, щоб перевірити, чи він уже існує, оскільки він повертає `true`, коли масив уже містить це значення.

Створіть функцію з назвою `addTag`, яка приймає параметри `id` і `tag`. Знайдіть артефакт за допомогою `collection[id]`. Якщо він існує і його масив `tags` ще не містить `tag`, додайте `tag` у `artifact.tags`.

# --hints--

У вас має бути функція з назвою `addTag` з параметрами `id` і `tag`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addTag);
const params = explorer.allFunctions.addTag.parameters;
assert.equal(params.length, 2);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "tag");
```

Виклик `addTag` з новим тегом має додати цей тег у масив `tags` артефакта.

```js
addTag(102, "ancient");
assert.include(collection[102].tags, "ancient");
```

Виклик `addTag` з тегом, який артефакт уже має, не повинен додавати дублікати.

```js
const tagsBefore = [...collection[102].tags];
addTag(102, "bronze");
assert.deepEqual(collection[102].tags, tagsBefore);

const start = collection[102].tags.length;
addTag(102, "unique-check");
addTag(102, "unique-check");
assert.strictEqual(collection[102].tags.length, start + 1);
```

`addTag` має додавати тег до того артефакта, на який посилається `id`.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
addTag(457, "test-tag");
assert.deepEqual(collection[457].tags, ["test-tag"]);
```

`addTag` не повинен викликати помилку або змінювати колекцію, якщо жоден артефакт не відповідає `id`.

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => addTag(999, "test"));
assert.strictEqual(JSON.stringify(collection), before);
```

Ви маєте додавати `tag` у існуючий масив `tags` артефакта, а не замінювати цей масив.

```js
collection[458] = {
  title: "Push Test",
  category: "Test Category",
  curator: { id: 458, name: "Test Curator" },
  locations: [{ gallery: "Hall Z", year: 2002 }],
  tags: [],
  onDisplay: false,
};
const tagsBefore = collection[458].tags;
const spy = __helpers.spyOn(tagsBefore, "push");
try {
  addTag(458, "pushed");
  assert.deepEqual(spy.calls, [["pushed"]]);
} finally {
  spy.restore();
}
assert.strictEqual(collection[458].tags, tagsBefore);
assert.deepEqual(collection[458].tags, ["pushed"]);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

--fcc-editable-region--

--fcc-editable-region--
```
