---
id: 6a57c331ae94fa4631b60866
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Іноді співробітники просять артефакт, якого немає в музеї. Якщо під цим `id` не збережено жодного артефакту, `collection[id]` буде `undefined`, і спроба прочитати `.title` викличе помилку.

Оновіть оператор return, щоб використати умовний вираз. Поверніть `artifact.title`, якщо артефакт знайдено, і рядок `"Artifact not found"`, якщо ні.

# --hints--

`getArtifactTitle(101)` має повертати `"Golden Mask"`.

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

`getArtifactTitle` має повертати `"Artifact not found"`, якщо жоден артефакт не відповідає `id`.

```js
assert.strictEqual(getArtifactTitle(999), "Artifact not found");
```

`getArtifactTitle` має повертати назву будь-якого знайденого артефакту, а не лише `101`.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
assert.strictEqual(getArtifactTitle(457), "Test Artifact");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
--fcc-editable-region--
  return artifact.title;
--fcc-editable-region--
}
```
