---
id: 6a57c331ae94fa4631b60869
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

Нові дослідження підтверджують, що Золота маска належала королівській родині. Додайте тег `"royal"` до артефакту з `id` `101`. Потім викличте `console.log()` з `collection[101].tags`, щоб вивести теги цього артефакту.

# --hints--

Ви повинні викликати `addTag` з `101` і `"royal"`.

```js
assert.match(__helpers.removeJSComments(code), /addTag\s*\(\s*101\s*,\s*["']royal["']\s*\)/);
```

Ви повинні вивести в консоль `tags` першого артефакту.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*tags|\[\s*["']tags["']\s*\])\s*\)/);
```

Ви повинні вивести `tags` після додавання нового тегу.

```js
assert.match(__helpers.removeJSComments(code), /addTag\s*\(\s*101\s*,\s*["']royal["']\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*tags|\[\s*["']tags["']\s*\])\s*\)/);
```

У `tags` першого артефакту має бути `"royal"`.

```js
assert.include(collection[101].tags, "royal");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
