---
id: 6a57c265abbc95dfebc05098
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви створите аудитор походження артефактів. Ви збережете музейні артефакти в об’єкті та напишете функції для їх пошуку, тегування, переміщення між галереями та підсумовування їхньої історії.

Кожен артефакт зберігається під своїм ідентифікатором. Створіть об’єкт з назвою `collection` і додайте до нього запис. Встановіть ключ `101`, а значення — об’єкт із `title` зі значенням `"Golden Mask"` та `category` зі значенням `"Ceremonial"`.

# --hints--

Ви повинні оголосити змінну з назвою `collection`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.collection);
```

`collection` має бути об’єктом.

```js
assert.isObject(collection);
```

У `collection` має бути один запис.

```js
assert.strictEqual(Object.keys(collection).length, 1);
```

Артефакт, збережений під `101`, має бути об’єктом.

```js
assert.isObject(collection[101]);
```

Артефакт, збережений під `101`, має мати `title` зі значенням `"Golden Mask"`.

```js
assert.strictEqual(collection[101].title, "Golden Mask");
```

Артефакт, збережений під `101`, має мати `category` зі значенням `"Ceremonial"`.

```js
assert.strictEqual(collection[101].category, "Ceremonial");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
