---
id: 6a57c331ae94fa4631b60864
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Аудит починається з читання того, що вже зафіксовано. Оскільки кожен артефакт зберігається під своїм ідентифікатором, ви можете отримати до нього доступ, помістивши цей ідентифікатор у квадратні дужки, наприклад `collection[101]`.

Викличте `console.log()` з `collection[101].title`, щоб вивести в консоль назву першого артефакту. Потім викличте `console.log()` з `collection[101].curator.name`, щоб вивести ім’я його куратора.

# --hints--

Ви повинні вивести в консоль `title` першого артефакту.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*title|\[\s*["']title["']\s*\])\s*\)/);
```

Ви повинні вивести в консоль `name` куратора першого артефакту.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

Спочатку виведіть `title`, а потім `name` куратора.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*title|\[\s*["']title["']\s*\])\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

--fcc-editable-region--

--fcc-editable-region--
```
