---
id: 6a57c331ae94fa4631b60861
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Кожен артефакт має відповідальну за нього особу.

Додайте властивість `curator` до артефакта, що зберігається під ключем `101`, і встановіть її значення як об’єкт з `id` `201` та `name` `"Earl Sinclair"`.

# --hints--

Артефакт, що зберігається під `101`, має мати властивість `curator`.

```js
assert.property(collection[101], "curator");
```

`curator` має бути об’єктом.

```js
assert.isObject(collection[101].curator);
```

`curator` повинен мати `id` зі значенням `201`.

```js
assert.strictEqual(collection[101].curator.id, 201);
```

`curator` має мати властивість `name` зі значенням `"Earl Sinclair"`.

```js
assert.strictEqual(collection[101].curator.name, "Earl Sinclair");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
