---
id: 6a57c331ae94fa4631b60863
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

У музеї є більше ніж один артефакт у записах. Додайте другий артефакт до `collection` під ключем `102`, дотримуючись тієї ж форми, що й у першого.

Використовуйте ці значення:

```js
{
  title: "Bronze Tablet",
  category: "Inscription",
  curator: {
    id: 202,
    name: "Robert Sinclair",
  },
  locations: [{ gallery: "Archive Wing", year: 2019 }],
  tags: ["bronze", "writing"],
  onDisplay: false,
}
```

# --hints--

У `collection` має бути два записи.

```js
assert.strictEqual(Object.keys(collection).length, 2);
```

Артефакт, збережений під ключем `102`, має відповідати описаному об’єкту.

```js
assert.deepEqual(collection[102], {
  title: "Bronze Tablet",
  category: "Inscription",
  curator: {
    id: 202,
    name: "Robert Sinclair",
  },
  locations: [{ gallery: "Archive Wing", year: 2019 }],
  tags: ["bronze", "writing"],
  onDisplay: false,
});
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
--fcc-editable-region--
  
--fcc-editable-region--
};
```
