---
id: 66adc82170acb1464bd348eb
title: Paso 13
challengeType: 1
dashedName: step-13
---

# --description--

Los últimos mensajes del bot se centrarán en el tema favorito del bot.

Comienza creando una nueva variable llamada `favoriteSubject` y asígnale la cadena `"Computer Science"`.

# --hints--

Debes tener una variable llamada `favoriteSubject`.

```js
assert.isNotNull(favoriteSubject);
```

Tu variable `favoriteSubject` debe ser una cadena.

```js
assert.isString(favoriteSubject);
```

Debes asignar la cadena `"Computer Science"` a la variable `favoriteSubject`.

```js
assert.strictEqual(favoriteSubject, "Computer Science");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

--fcc-editable-region--

--fcc-editable-region--
```
