---
id: 66adc42868cab843ccee87d9
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

Ahora es momento de ver el nuevo valor de `bot`.

Empiece por crear una nueva variable llamada `nicknameIntroduction`.

Use la concatenación de cadenas para unir la cadena `"My nickname is "` con la variable `bot` seguida de un punto (`.`).

Asigne la cadena resultante a la variable `nicknameIntroduction`.

Luego, registre el valor de `nicknameIntroduction` en la consola.

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

Debe tener una variable llamada `nicknameIntroduction`.

```js
assert.isNotNull(nicknameIntroduction);
```

Su variable `nicknameIntroduction` debe ser una cadena.

```js
assert.isString(nicknameIntroduction);
```

Debe usar la concatenación de cadenas para unir la cadena `"My nickname is "` con la variable `bot` seguida de un punto (`.`). Ten en cuenta los espacios.

```js
assert.strictEqual(nicknameIntroduction, "My nickname is professorBot.");
```

Debes tener una instrucción `console` en tu código.

```js
assert.lengthOf(logSpy.calls, 6);
```

Debe registrar el valor de `nicknameIntroduction` en la consola.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*nicknameIntroduction\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

--fcc-editable-region--

--fcc-editable-region--
```
