---
id: 66ada3f46945763dd97f43f8
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Anteriormente, creaste las variables `bot` y `botLocation`. Ahora, las usarás para mostrar nuevos mensajes en la consola.

En lecciones anteriores, aprendiste cómo trabajar con la concatenación de cadenas usando el operador `+` para concatenar cadenas así:

```js
let firstName = "John";

console.log("Hello, my name is " + firstName + ".");
// result: "Hello, my name is John."
```

Recuerda que debes tener cuidado con los espacios al concatenar cadenas con variables.

Crea una variable llamada `botIntroduction`.

Luego, usa la concatenación de cadenas con el operador `+` para unir la cadena `"My name is "` seguido de la variable `bot` y seguido de un punto (`.`). 

Asigna este valor a la variable `botIntroduction`.

Luego, registra la variable `botIntroduction` en la consola.

# --hints--

Debes tener una variable llamada `botIntroduction`.

```js
assert.isNotNull(botIntroduction);
```

Tu variable `botIntroduction` debe contener el valor de una cadena.

```js
assert.isString(botIntroduction);
```

Debes usar concatenación de cadenas con el operador `+` para unir la cadena `"My name is "` con la variable `bot`, seguido de un punto (`.`). Ten cuidado con los espacios.  Assign this concatenated string to the `botIntroduction` variable.

```js
assert.equal(botIntroduction, "My name is teacherBot.");
```

Debes tener una cuarta declaración `console.log()` en tu código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 4);
```

Tu declaración de `console` debe mostrar la variable `botIntroduction`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botIntroduction\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

--fcc-editable-region--

--fcc-editable-region--
```
