---
id: 66ad9c633fa26d3c1475eae3
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

Ahora, es momento de añadir otro mensaje del bot.

Añade otra declaración de `console` al código que registre el mensaje `"Allow me to introduce myself."`.

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

Debe haber una tercera declaración de `console.log()` en tu código.

```js
assert.lengthOf(logSpy.calls, 3);
```

Tu declaración de `console` debería tener el mensaje `"Allow me to introduce myself."`. No olvides las comillas.

```js
assert.deepInclude(logSpy.calls, ["Allow me to introduce myself."]);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
