---
id: 66ad8294a0ad902f1b31b612
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, aprenderás a trabajar con los fundamentos de JavaScript construyendo un bot de saludo.

En este primer paso, querrás mostrar un mensaje en la consola desde el bot de saludo.

Recuerda que aprendiste sobre `console.log()` y cadenas en lecciones anteriores.

Aquí tienes un recordatorio de cómo usar `console.log()` con cadenas de texto:

```js
console.log("Hello, World!");
```

Agrega una declaración `console.log()` que muestre la cadena `"Hi there!"` en la consola. ¡No olvides tus comillas alrededor del mensaje!

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

Debes tener un `console.log()` en tu código.

```js
assert.isNotEmpty(logSpy.calls);
```

Tu declaración `console` debe tener el mensaje `"Hi there!"`. No olvides las comillas.

```js
assert.deepInclude(logSpy.calls, ["Hi there!"]);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
