---
id: 66adc42868cab843ccee87d9
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

Now it is time to see the new `bot` value.

Empiece por crear una nueva variable llamada `nicknameIntroduction`.

Use la concatenaci贸n de cadenas para unir la cadena `"Mi apodo es "` con la variable `bot` seguida de un punto (`.`).

Asigne la cadena resultante a la variable `nicknameIntroduction`.

Luego, registre el valor de `nicknameIntroduction` en la consola. 

# --hints--

Debe tener una variable llamada `nicknameIntroduction`.

```js
assert.isNotNull(nicknameIntroduction);
```

Su variable `nicknameIntroduction` debe ser una cadena.

```js
assert.isString(nicknameIntroduction);
```

Debe usar la concatenaci贸n de cadenas para unir la cadena `"Mi apodo es "` con la variable `bot` seguida de un punto (`.`). Ten en cuenta los espacios.

```js
assert.strictEqual(nicknameIntroduction, "My nickname is professorBot.");
```

Debes tener una instrucci贸n `console` en tu c贸digo.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 6);
```

Debe registrar el valor de `nicknameIntroduction` en la consola.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*nicknameIntroduction\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

--fcc-editable-region--

--fcc-editable-region--
```
