---
id: 66ad8d150264db3926eccfeb
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

En lecciones anteriores, aprendiste sobre la palabra clave `let` y cómo declarar variables en JavaScript.

Aquí tienes un recordatorio de cómo declarar una variable usando la palabra clave `let`:

```js
let greeting;
```

Usa la palabra clave `let` para declarar una variable llamada `bot`.

**NOTA**: Estás usando `let` aquí porque más adelante en el taller, estarás cambiando el valor de la variable `bot`.

# --hints--

Debes tener la palabra clave `let` en tu código.

```js
assert.match(__helpers.removeJSComments(code), /let/g);
```

Debes declarar una variable llamada `bot`.

```js
assert.isNotNull(bot);
```

Debes usar la palabra clave `let` para declarar la variable.

```js
const explorer = await __helpers.Explorer(code);
const { bot } = explorer.variables;
assert.isTrue(bot?.matches('let bot'));
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
--fcc-editable-region--

--fcc-editable-region--
```
