---
id: 66ad984e92ce213a66f7e51c
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Ahora, es momento de asignar un valor a tu variable `bot`.

En lecciones anteriores, aprendiste cómo asignar valores a variables así:

```js
variableName = "Here is the value";
```

Recuerda que lo que está en el lado derecho del operador de asignación `=` es el valor que estás asignando a la variable en el lado izquierdo.

Asigna la cadena `"teacherBot"` a la variable `bot`.

# --hints--

Debes asignar una cadena a tu variable `bot`.

```js
assert.isString(bot);
```

Tu variable `bot` debe tener asignada la cadena `"teacherBot"`.

```js
assert.equal(bot, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
