---
id: 66adb844118ba74107ce771f
title: Paso 8
challengeType: 1
dashedName: step-8
---

# --description--

El siguiente mensaje del bot se referirá a la ubicación del bot.

Cree una variable llamada `botLocationSentence`.

Luego, use la concatenación de cadenas con el operador `+` para unir la cadena `"I live in "` con la variable `botLocation` seguida de un punto (`.`).

Asigne este valor a la variable `botLocationSentence`.

Luego, registre el valor de `botLocationSentence` en la consola.

# --hints--

Debe tener una variable llamada `botLocationSentence`.

```js
assert.isNotNull(botLocationSentence);
```

Su variable `botLocationSentence` debe contener el valor de una cadena.

```js
assert.isString(botLocationSentence);
```

Debe usar la concatenación de cadenas con el operador `+` para unir la cadena `"I live in "` con la variable `botLocation` seguida de un punto (`.`). Tenga cuidado con los espacios.

```js
assert.equal(botLocationSentence, "I live in the universe.");
```

Debe tener una sentencia `console` en su código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 5);
```

Su sentencia `console` debe mostrar la variable `botLocationSentence`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botLocationSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
