---
id: 66ad984e92ce213a66f7e51c
title: Pasos 4
challengeType: 1
dashedName: step-4
---

# --description--

Cuando necesitas declarar variables con múltiples palabras, puedes usar la convención de nomenclatura <dfn>camelCase</dfn>.

Cuando usas `camelCase`, la primera palabra está en minúscula y la primera letra de cada palabra subsiguiente está capitalizada.

Aquí hay un ejemplo:

```js
let thisIsCamelCase;
```

Usa `let` para declarar una variable llamada `botLocation`.

# --hints--

Deberías usar `let` para declarar la variable.

```js
assert.lengthOf(code.match(/let/g), 2);
```

Tu variable debería llamarse `botLocation`.

```js
assert.isNotNull(botLocation);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
