---
id: 66ad8ab945d266383f318cbf
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Ahora deberías ver el primer mensaje del bot en la consola.

Es hora de añadir un segundo mensaje del bot.

Agrega otra declaración `console.log` para mostrar el mensaje `"I am excited to talk to you."`.

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

Deberías tener una segunda instrucción `console.log()` en tu código.

```js
assert.lengthOf(logSpy.calls, 2);
```

Tu declaración `console` debe tener el mensaje `"I am excited to talk to you."`. No olvides las comillas.

```js
assert.deepInclude(logSpy.calls, ["I am excited to talk to you."]);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
--fcc-editable-region--

--fcc-editable-region--
```
