---
id: 66ad998285df023b23dacdd3
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Ahora, es momento de inicializar la variable `botLocation`.

Cuando necesitas declarar variables con múltiples palabras, puedes usar la convención de nomenclatura <dfn>camelCase</dfn>.

Cuando usas `camelCase`, la primera palabra está en minúscula y la primera letra de cada palabra subsiguiente está capitalizada. 

Aquí hay un ejemplo:

```js
let thisIsCamelCase = "Alice the camel is actually a horse.";
```

Declara y asigna la cadena `"the universe"` a la variable `botLocation` en la misma línea usando la palabra clave `let`.

# --hints--

Usa la palabra clave `let` para declarar una variable.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

Tu variable debería llamarse `botLocation`.

```js
assert.isNotNull(botLocation);
```

Debes asignar una cadena a tu variable `botLocation`.

```js
assert.isString(botLocation);
```

Tu variable `botLocation` debe tener asignada la cadena `"the universe"`.

```js
assert.equal(botLocation, "the universe");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
