---
id: 669fde2081f65141ad703fe4
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

ページの最初のセクションは about セクションになります。このセクションでは Mr. Whiskers を紹介し、このブログがどんな内容かユーザーに伝えます。

`main` 要素の中に、`section` 属性が `id` に設定された `"about"` 要素を追加してください。

`section` 要素の中に、テキストが `h2` の `About` を追加してください。

# --hints--

`section` 要素には開始タグが必要です。

```js
assert.match(code, /<section.*>/);
```

`section` 要素には終了タグが必要です。

```js
assert.match(code, /<\/section\>/);
```

`section` 要素には `id` が `"about"` に設定されている必要があります。

```js
const sectionElement = document.querySelector('section');
assert.strictEqual(sectionElement?.getAttribute("id"), "about");
```

`section` 要素は `main` 要素の中にある必要があります。

```js
const sectionElement = document.getElementById('about');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

`h2` 要素には開始タグが必要です。

```js
assert.match(code, /<h2\>/);
```

`h2` 要素には終了タグが必要です。

```js
assert.match(code, /<\/h2\>/);
```

`h2` 要素には `About` のテキストが必要です。スペルをもう一度確認してください。

```js
const h2Element = document.querySelector('body h2');
assert.strictEqual(h2Element?.innerText.trim(), "About");
```

`h2` 要素は `section` 要素の中にある必要があります。

```js
const h2Element = document.querySelector('#about h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
