---
id: 669fdfc9a5e3da42d2376609
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

about セクションを追加したので、`About` のリンクをクリックして、そのセクションにページがジャンプするか試してください。

ブログページの次のセクションは、Mr. Whiskers について書かれた投稿のリストになります。

`section` 要素をもう一つ追加し、`id` を `"posts"` に設定してください。

`section` 要素の中に、テキストが `h2` の `Posts` 要素を追加してください。

# --hints--

コードには開く `section` タグが必要です。

```js
assert.match(code, /<section.*>/i);
```

コードには閉じる `section` タグが必要です。

```js
assert.match(code, /<\/section>/i);
```

`section` 要素には `id` が `"posts"` に設定されている必要があります。

```js
assert.match(code, /<section.*id="posts".*>/i);
```

`section` 要素は `main` 要素の中にある必要があります。

```js
const sectionElement = document.getElementById('posts');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

`h2` 要素は `section` 要素の中にある必要があります。

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

`h2` 要素には `Posts` のテキストが必要です。

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.innerText.trim(), 'Posts');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </main>
  </body>
</html>
```
