---
id: 669fcb5106eeb13ab7afcd09
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

このワークショップでは、セマンティックHTMLを使って、猫のMr. Whiskersに関するブログページを作成する練習をします。

最初に作成するのはページの<dfn>ヘッダー</dfn>セクションです。

`header`要素は、ページのナビゲーションやその他の導入情報のような紹介コンテンツを表すために使われます。

以下は`header`要素を使った例です。

```html
<header>
  <h1>Main Page Title Goes Here</h1>
  <img src="example-logo.png" alt="Example logo" />
</header>
```

`body`要素の中に`header`要素を追加してください。

# --hints--

開く`header`タグがあるはずです。

```js
assert.match(code, /<header>/i);
```

閉じる`header`タグがあるはずです。

```js
assert.match(code, /<\/header>/i);
```

`header`要素は`body`要素の中にあるはずです。

```js
assert.exists(document.querySelector('body > header'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </body>
</html>
```
