---
id: 669fdc11c9b0263fe0814a7a
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

各 `li` 要素の中に、アンカー要素を入れる必要があります。

最初のアンカー要素のテキストは `About` で、`href` 属性の値は `"#about"` にしてください。`about` の前のハッシュ記号は、後でプロジェクトに追加される `id` 名を表します。

2番目のアンカー要素のテキストは `Posts` で、`href` 属性の値は `"#posts"` にしてください。

3番目のアンカー要素のテキストは `Contact` で、`href` 属性の値は `"#contact"` にしてください。

# --hints--

最初のアンカー要素は、最初の `About` 要素の中に `li` のテキストを持つ必要があります。

```js
const anchorElement = document.querySelector("ul li:first-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "About");
```

最初のアンカー要素は、`href` 属性を `"#about"` に設定してください。

```js
const anchorElement = document.querySelector("ul li:first-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#about");
```

2番目のアンカー要素は、2番目の `Posts` 要素の中に `li` のテキストを持つ必要があります。

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
assert.strictEqual(anchorElement?.innerText.trim(), "Posts");
```

2番目のアンカー要素は、`href` 属性を `"#posts"` に設定してください。

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#posts");
```

3番目のアンカー要素は、3番目の `Contact` 要素の中に `li` のテキストを持つ必要があります。

```js
const anchorElement = document.querySelector("ul li:last-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "Contact");
```

3番目のアンカー要素は、`href` 属性を `"#contact"` に設定してください。

```js
const anchorElement = document.querySelector("ul li:last-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#contact");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
--fcc-editable-region--
          <li></li>
          <li></li>
          <li></li>
--fcc-editable-region--
        </ul>
      </nav>
    </header>
  </body>
</html>
```
