---
id: 669fdd3965f36f40af9615e5
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

ページのヘッダーの作成が終わったので、メインコンテンツの追加を始めましょう。

`header` の下に、`main` 要素を追加してください。

# --hints--

`main` 要素には開始タグが必要です。

```js
assert.match(code, /<main\>/);
```

`main` 要素には終了タグが必要です。

```js
assert.match(code, /<\/main\>/);
```

`main` 要素は `header` 要素の下に配置してください。

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
