---
id: 669fd003cb89ee3c2402e041
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

この導入コンテンツでは、Mr. Whiskers の画像とキャプションを表示したいです。

`h1` 要素の下に、まず `figure` 要素を追加してください。

`figure` 要素の中に、`img` 要素を追加してください。

`src` の `img` 属性は `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"` の値を持ち、`alt` のテキストは `"a cat in the garden"` の値を持つべきです。

`img` 要素の下に、テキストが `figcaption` の `Mr. Whiskers in the Garden` を追加してください。

# --hints--

`figure` 要素は開始タグを持つべきです。

```js
assert.match(code, /<figure\>/);
```

`figure` 要素は終了タグを持つべきです。

```js
assert.match(code, /<\/figure\>/);
```

`figure` 要素は `header` 要素の中にあるべきです。

```js
const figureElement = document.querySelector('figure');
assert.equal(figureElement?.parentElement.tagName, 'HEADER');
```

`img` 要素が必要です。

```js
assert.exists(document.querySelector("img"));
```

`img` 要素は `figure` 要素の中にあるべきです。

```js
const imgElement = document.querySelector('img');
assert.strictEqual(imgElement?.parentElement.tagName, 'FIGURE');
```

`img` 要素は `src` 属性を `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"` に設定するべきです。

```js
const imgElement = document.querySelector('img');
const src = imgElement?.getAttribute('src');

assert.strictEqual(src, 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

`img` 要素は `alt` 属性を `"a cat in the garden"` に設定するべきです。

```js
const imgElement = document.querySelector('img');
const alt = imgElement?.getAttribute('alt');

assert.strictEqual(alt, 'a cat in the garden');
```

`figcaption` 要素は開始タグを持つべきです。

```js
assert.match(code, /<figcaption\>/);
```

`figcaption` 要素は終了タグを持つべきです。

```js
assert.match(code, /<\/figcaption\>/);
```

`figcaption` 要素は `figure` 要素の中にあるべきです。

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.parentElement.tagName, 'FIGURE');
```

`figcaption` 要素は `Mr. Whiskers in the Garden` のテキストを持つべきです。

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.innerText.trim(), "Mr. Whiskers in the Garden")
```

`figcaption` は `img` 要素の下にあるべきです。順番が間違っています。

```js
const figureElement = document.querySelector('figure');
assert.strictEqual(figureElement?.lastElementChild?.tagName, "FIGCAPTION")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
