---
id: 669fd6fd12918e3de87854d4
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

ブログには、ユーザーがページ内の異なるセクションに移動できるナビゲーションが必要です。

`nav` 要素は、ドキュメント内の他のセクションやウェブサイト内の他のセクションへのナビゲーションリンクを提供するために使われます。`nav` 要素はメニューや目次によく使われるのを見かけます。

`nav` 要素の使用例を示します：

```html
<nav>
  <ul>
    <li><a href="#">Home</a></li>
    <li><a href="#about">About</a></li>
    <li><a href="#contact">Contact</a></li>
  </ul>
</nav>
```

`figure` 要素の下に、`nav` 要素を追加し、その中にネストされた `ul` 要素を入れてください。

`ul` 要素の中に、3つの `li` 要素を追加してください。

# --hints--

`nav` 要素には開始タグが必要です。

```js
assert.match(code, /<nav\>/);
```

`nav` 要素には終了タグが必要です。

```js
assert.match(code, /<\/nav\>/);
```

`nav` 要素は `header` 要素の中にある必要があります。

```js
const navElement = document.querySelector('nav');
assert.strictEqual(navElement?.parentElement.tagName, 'HEADER');
```

`ul` 要素には開始タグが必要です。

```js
assert.match(code, /<ul\>/);
```

`ul` 要素には終了タグが必要です。

```js
assert.match(code, /<\/ul\>/);
```

`ul` 要素は `nav` 要素の中にある必要があります。

```js
const ulElement = document.querySelector('ul');
assert.strictEqual(ulElement?.parentElement.tagName, 'NAV');
```

`ul` 要素には3つの `li` 要素が含まれている必要があります。

```js
const ulElement = document.querySelector("ul");

const liElements = ulElement?.querySelectorAll("li");
assert.strictEqual(liElements?.length, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
