---
id: 66a335e5888ffb367633200a
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

このブログ記事には、<dfn>lorem ipsum</dfn> テキストを使ったいくつかの段落が含まれます。

Lorem ipsum はウェブ開発でよく使われるプレースホルダーテキストです。ウェブページの基本構造の作成に集中したいときに、実際のコンテンツを気にせずに済むので便利です。

以下は lorem ipsum を使った例です。

```html
<p>
  Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
  quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
  quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
  quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
</p>
```

`h3` 要素の下に、lorem ipsum テキストの段落を2つ追加してください。

# --hints--

`article` 要素の中に段落要素を2つ入れているはずです。

```js
const paragraphElements = document.querySelectorAll('#posts article p');
assert.strictEqual(paragraphElements.length, 2);
```

最初の段落要素には lorem ipsum テキストを入れてください。スペルや句読点の間違いがないかよく確認しましょう。

```js
const paragraphElement = document.querySelector('article p:first-of-type');
const loremIpsum = `Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.`

assert.strictEqual(paragraphElement?.innerText.trim().replaceAll(/\s+/g, ' '), loremIpsum);
```

2つ目の段落要素にも lorem ipsum テキストを入れてください。スペルや句読点の間違いがないかよく確認しましょう。

```js
const paragraphElement = document.querySelector('article p:last-of-type');
const loremIpsum = `Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.`

assert.strictEqual(paragraphElement?.innerText.trim().replaceAll(/\s+/g, ' '), loremIpsum);
```

段落は `h3` 要素の後に配置してください。

```js
let h3 = document.querySelectorAll('article > p ~ h3');
assert.isEmpty(h3);
const paragraphs = document.querySelectorAll('article > h3 ~ p');
assert.lengthOf(paragraphs, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>
        <article>
          <h3>Mr. Whiskers' First Day Home</h3>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </article>
      </section>
    </main>
  </body>
</html>
```
