---
id: 66a334a1a7cca6354999f9bf
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

最初のブログ投稿には、<dfn>article</dfn> 要素を使います。

`article` 要素はウェブページ上の自己完結型のコンテンツを表します。

```html
<article>
  <h1>Example heading</h1>
  <p>Example article text</p>
</article>
```

`h2` 要素の下に、`article` 要素を追加してください。

`article` 要素の中に、テキスト `h3` を持つ `Mr. Whiskers' First Day Home` 要素を追加してください。

ここで `h3` を使う理由は、見出し要素の適切な構造階層を維持することが重要だからです。投稿のサブ見出しが `h2` 要素なので、階層の次のレベルは `h3` になります。

# --hints--

コードには開く `article` タグが必要です。

```js
assert.match(code, /<article>/i);
```

コードには閉じる `article` タグが必要です。

```js
assert.match(code, /<\/article>/i);
```

`article` 要素は `section` 要素の中にあり、`id` は `posts` に設定されている必要があります。

```js
const articleElement = document.querySelector('#posts article');
assert.strictEqual(articleElement?.parentElement.tagName, 'SECTION');
```

コードには開く `h3` タグが必要です。

```js
assert.match(code, /<h3>/i);
```

コードには閉じる `h3` タグが必要です。

```js
assert.match(code, /<\/h3>/i);
```

`h3` 要素には `Mr. Whiskers' First Day Home` のテキストが必要です。

```js
const h3Element = document.querySelector('h3');
assert.strictEqual(h3Element?.innerText.trim(), "Mr. Whiskers' First Day Home");
```

`h3` 要素は `article` 要素の中にある必要があります。

```js
const h3Element = document.querySelector('#posts article h3');
assert.strictEqual(h3Element?.parentElement.tagName, 'ARTICLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
