---
id: 66a49f685961e997e337cab1
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`header` はブログのメインタイトル、画像、ページナビゲーションを表示する役割を担います。

`header` 要素の中に、`h1` のテキストを持つ `Welcome to Mr. Whiskers' Blog Page!` を追加してください。

# --hints--

開く `h1` タグがあるはずです。

```js
assert.match(code, /<h1>/i);
```

閉じる `h1` タグがあるはずです。

```js
assert.match(code, /<\/h1>/i);
```

`h1` 要素には `Welcome to Mr. Whiskers' Blog Page!` のテキストがあるはずです。スペルミスがないかもう一度確認してください。

```js
const h1Text = document.querySelector("h1")?.innerText.trim();
assert.strictEqual(h1Text, "Welcome to Mr. Whiskers' Blog Page!");
```

`h1` 要素は `header` 要素の中にネストされているはずです。

```js
const h1Element = document.querySelector("h1");
assert.strictEqual(h1Element?.parentElement.tagName, 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
