---
id: 66a7e72adf226c02626715a3
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

`h2` 要素の下に、`Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.` のテキストを持つ段落要素を追加してください。

段落要素の下に、`His playful nature and boundless energy keeps me on my toes. I love him so much.` のテキストを持つ別の段落要素を追加してください。

# --hints--

最初の段落には `Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.` のテキストを入れてください。スペルをよく確認しましょう。

```js
const pElement = document.querySelector('body p:first-of-type');
const text = "Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers."
assert.strictEqual(pElement?.innerText.trim(), text);
```

最初の段落は `section` 要素の中にネストしてください。

```js
const pElement = document.querySelector('#about p:first-of-type');
assert.strictEqual(pElement?.parentElement.tagName, 'SECTION');
```

2つ目の段落は `section` 要素の中に入れてください。

```js
const pElement = document.querySelectorAll('#about p');
assert.strictEqual(pElement?.length, 2);
```

2つ目の段落には `His playful nature and boundless energy keeps me on my toes. I love him so much.` のテキストを入れてください。スペルをよく確認しましょう。

```js
const pElement = document.querySelector('#about p:last-of-type');
const text = "His playful nature and boundless energy keeps me on my toes. I love him so much."
assert.strictEqual(pElement?.innerText.trim(), text);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
