---
id: 587d78a5367417b2b2512ada
title: 使用 CSS Transform scale 屬性在暫留時縮放元素
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLPJuM'
forumTopicId: 301077
dashedName: use-the-css-transform-scale-property-to-scale-an-element-on-hover
---

# --description--

`transform` 屬性具有多種函式（程式），讓你可以縮放、移動、旋轉、傾斜等你的元素。當與指定元素特定狀態的偽類（例如 `:hover`）一起使用時，`transform` 屬性可以輕鬆地為你的元素添加互動性。

以下是一個範例，當使用者暫留在段落元素上時，將其縮放為原始大小的 2.1 倍：

```css
p:hover {
  transform: scale(2.1);
}
```

**注意：** 對 `div` 元素套用轉換也會影響該 div 中包含的任何子元素。

# --instructions--

為你的 `div` 的 `hover` 狀態添加一條 CSS 規則，並使用 `transform` 屬性，當使用者暫留其上方時，將 `div` 元素縮放到原始大小的 1.1 倍。

# --hints--

當使用者暫留在 `div` 元素上時，其大小應該放大 1.1 倍。

```js
assert.match(code, /div:hover\s*?{\s*?transform:\s*?scale\(1\.1\);/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
    background: linear-gradient(
      53deg,
      #ccfffc,
      #ffcccf
    );
  }



</style>

<div></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
    background: linear-gradient(
      53deg,
      #ccfffc,
      #ffcccf
    );
  }
  div:hover {
    transform: scale(1.1);
  }
</style>
<div></div>
```
