---
id: 587d78a3367417b2b2512ad0
title: 使用 margin 屬性水平置中元素
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLJqU4'
forumTopicId: 301043
dashedName: center-an-element-horizontally-using-the-margin-property
---

# --description--

另一種定位技術是將區塊元素水平置中。實現此目的的一種方法是將其 `margin` 設為 auto。

這個方法也適用於影像。影像預設是行內元素，但當你將 `display` 屬性設定為 `block` 時，可以改成區塊元素。

# --instructions--

透過為 `div` 添加值為 `auto` 的 `margin` 屬性，將其置中於頁面。

# --hints--

`div` 應該設定 `margin` 為 `auto`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('div')?.margin ,'auto');
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    background-color: blue;
    height: 100px;
    width: 100px;

  }
</style>
<div></div>
```

# --solutions--

```html
<style>
  div {
    background-color: blue;
    height: 100px;
    width: 100px;
    margin: auto;
  }
</style>
<div></div>
```
