---
id: 587d78a8367417b2b2512ae3
title: 使用無限動畫次數持續動畫元素
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJDVfq'
forumTopicId: 301041
dashedName: animate-elements-continually-using-an-infinite-animation-count
---

# --description--

先前的挑戰涵蓋了如何使用部分動畫屬性和 `@keyframes` 規則。另一個動畫屬性是 `animation-iteration-count`，它允許你控制動畫要重複播放多少次。以下是一個範例：

```css
animation-iteration-count: 3;
```

在這種情況下，動畫會在執行 3 次後停止，但可以透過將該值設定為 `infinite` 來讓動畫持續執行。

# --instructions--

為了讓球在右側持續不斷地彈跳，將 `animation-iteration-count` 屬性更改為 `infinite`。

# --hints--

`animation-iteration-count` 屬性應該有 `infinite` 的值。

```js
const ballElement = document.querySelector('#ball');
const ballStyle = window.getComputedStyle(ballElement); 
assert.equal(ballStyle?.animationIterationCount, 'infinite');
```

# --seed--

## --seed-contents--

```html
<style>

  #ball {
    width: 100px;
    height: 100px;
    margin: 50px auto;
    position: relative;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: bounce;
    animation-duration: 1s;
    animation-iteration-count: 3;
  }

  @keyframes bounce{
    0% {
      top: 0px;
    }
    50% {
      top: 249px;
      width: 130px;
      height: 70px;
    }
    100% {
      top: 0px;
    }
  }
</style>
<div id="ball"></div>
```

# --solutions--

```html
<style>
  #ball {
    width: 100px;
    height: 100px;
    margin: 50px auto;
    position: relative;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: bounce;
    animation-duration: 1s;
    animation-iteration-count: infinite;
  }

  @keyframes bounce{
    0% {
      top: 0px;
    }
    50% {
      top: 249px;
      width: 130px;
      height: 70px;
    }
    100% {
      top: 0px;
    }
  }
</style>
<div id="ball"></div>
```
