---
id: 587d78a9367417b2b2512aea
title: 使用貝茲曲線讓動作更自然
challengeType: 0
videoUrl: 'https://scrimba.com/c/c7akWUv'
forumTopicId: 301063
dashedName: make-motion-more-natural-using-a-bezier-curve
---

# --description--

這個挑戰會為一個元素製作動畫，以複製球被拋接的動作。先前的挑戰涵蓋了 `linear` 和 `ease-out` 三次貝茲曲線，但兩者都無法準確描繪拋接的動作。你需要為此自訂一條貝茲曲線。

當 `animation-iteration-count` 設為 infinite 時，`animation-timing-function` 會在每個關鍵影格自動迴圈。由於在動畫持續時間中間（在 `50%`）設定了關鍵影格規則，因此球的向上和向下移動會產生兩個相同的動畫進度。

以下的三次貝茲曲線模擬了拋接動作：

```css
cubic-bezier(0.3, 0.4, 0.5, 1.6);
```

請注意 y2 的值大於 1。雖然三次貝茲曲線映射在 1×1 的座標系統上，且只能接受 0 到 1 之間的 x 值，但 y 值可以設定為大於 1 的數字。這會產生彈跳的動作，非常適合模擬拋接球。

# --instructions--

將 id 為 `green` 的元素的 `animation-timing-function` 值更改為 `cubic-bezier` 函式（程式），其 x1、y1、x2、y2 值分別設定為 0.311、0.441、0.444、1.649。

# --hints--

id 為 `green` 的元素的 `animation-timing-function` 屬性值應該是具有指定 x1、y1、x2、y2 值的 `cubic-bezier` 函式（程式）。

```js
const greenElement = document.querySelector('#green');
 const greenStyle = window.getComputedStyle(greenElement);
assert.equal(greenStyle?.animationTimingFunction, 'cubic-bezier(0.311, 0.441, 0.444, 1.649)');
```

# --seed--

## --seed-contents--

```html
<style>
  .balls {
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    top: 60%;
    animation-name: jump;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 25%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 50%;
    animation-timing-function: ease-out;
  }
  #green {
    background: green;
    left: 75%;
    animation-timing-function: cubic-bezier(0.69, 0.1, 1, 0.1);
  }

  @keyframes jump {
    50% {
      top: 10%;
    }
  }
</style>
<div class="balls" id="red"></div>
<div class="balls" id="blue"></div>
<div class="balls" id="green"></div>
```

# --solutions--

```html
<style>
  .balls {
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    top: 60%;
    animation-name: jump;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 25%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 50%;
    animation-timing-function: ease-out;
  }
  #green {
    background: green;
    left: 75%;
    animation-timing-function: cubic-bezier(0.311, 0.441, 0.444, 1.649);
  }

  @keyframes jump {
    50% {
      top: 10%;
    }
  }
</style>
<div class="balls" id="red"></div>
<div class="balls" id="blue"></div>
<div class="balls" id="green"></div>
```
