---
id: 587d781d367417b2b2512ac8
title: 調整錨點頁籤的暫留狀態
challengeType: 0
videoUrl: 'https://scrimba.com/c/cakRGcm'
forumTopicId: 301035
dashedName: adjust-the-hover-state-of-an-anchor-tag
---

# --description--

這個挑戰將涉及偽類的使用。偽類是一個可以添加到選擇器中的關鍵字，用來選擇元素的特定狀態。

例如，可以使用 `:hover` 偽類選擇器來更改錨點頁籤在暫留狀態下的樣式。以下是用於在暫留狀態時將錨點頁籤的 `color` 更改為紅色的 CSS：

```css
a:hover {
  color: red;
}
```

# --instructions--

程式碼編輯器有一個 CSS 規則將所有 `a` 頁籤的樣式設為黑色。為你的規則添加一條，當使用者暫留在 `a` 頁籤上時，`color` 為藍色。

# --hints--

錨點頁籤的 `color` 應保持為黑色，只為 `:hover` 狀態添加 CSS 規則。

```js
const anchorElement = document.querySelector("a"); 
const anchorStyle = window.getComputedStyle(anchorElement);
assert.equal(anchorStyle?.color, 'rgb(0, 0, 0)');
```

當暫留時，錨點頁籤的 `color` 應為藍色。

```js
assert.match(code,
    /a:hover\s*?{\s*?color:\s*?(blue|rgba\(\s*?0\s*?,\s*?0\s*?,\s*?255\s*?,\s*?1\s*?\)|#00F|rgb\(\s*?0\s*?,\s*?0\s*?,\s*?255\s*?\))\s*?;\s*?}/gi
  );
```

# --seed--

## --seed-contents--

```html
<style>
  a {
    color: #000;
  }



</style>
<a href="https://freecatphotoapp.com/" target="_blank">CatPhotoApp</a>
```

# --solutions--

```html
<style>
  a {
    color: #000;
  }
  a:hover {
    color: rgba(0,0,255,1);
  }
</style>
<a href="https://freecatphotoapp.com/" target="_blank">CatPhotoApp</a>
```
