---
id: 587d78a9367417b2b2512ae8
title: 了解貝茲曲線的運作方式
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bDrs8'
forumTopicId: 301058
dashedName: learn-how-bezier-curves-work
---

# --description--

上一個挑戰介紹了 `animation-timing-function` 屬性和一些在動畫持續時間內改變速度的關鍵字。CSS 提供了除了關鍵字之外的選項，可透過使用貝茲曲線來更精細地控制動畫的播放方式。

在 CSS 動畫中，貝茲曲線與 `cubic-bezier` 函式（程式）一起使用。曲線的形狀表述動畫的播放方式。該曲線存在於 1×1 的座標系統中。此座標系統的 X 軸是動畫的持續時間（可視為時間刻度），而 Y 軸是動畫的變化。

`cubic-bezier` 函式由四個主要點組成，這些點位於這個 1×1 的格線上：`p0`、`p1`、`p2` 和 `p3`。`p0` 和 `p3` 已為你設定－它們分別是起點和終點，始終位於原點 (0, 0) 和 (1, 1)。你設定另外兩個點的 x 和 y 值，而你將它們放置在格線上的位置決定動畫曲線的形狀。這是在 CSS 中透過宣告 `p1` 和 `p2`「錨點」的 x 和 y 值，以 `(x1, y1, x2, y2)` 的形式完成。綜合以上，以下是在 CSS 程式碼中使用貝茲曲線的範例：

```css
animation-timing-function: cubic-bezier(0.25, 0.25, 0.75, 0.75);
```

在上述範例中，每個點的 x 和 y 值是等同的（x1 = 0.25 = y1 和 x2 = 0.75 = y2），如果你還記得幾何類別，這會導致一條從原點延伸到點 (1, 1) 的直線。這個動畫是在動畫長度期間元素的線性變化，且與使用 `linear` 關鍵字相同。換句話說，它以恆定速度變化。

# --instructions--

對於 id 為 `ball1` 的元素，將 `animation-timing-function` 屬性的值從 `linear` 更改為其等價的 `cubic-bezier` 函式（程式）值。使用上述範例中給出的點值。

# --hints--

id 為 `ball1` 的元素的 `animation-timing-function` 屬性值應該是等價於線性的 `cubic-bezier` 函式（程式）。

```js
const ballOne = document.querySelector('#ball1');
const ballOneStyle = window.getComputedStyle(ballOne); 
assert.equal(
  ballOneStyle?.animationTimingFunction,
    'cubic-bezier(0.25, 0.25, 0.75, 0.75)'
);
```

id 為 `ball2` 的元素的 `animation-timing-function` 屬性值不應該改變。

```js
const ballTwoElement = document.querySelector('#ball2');
const ballTwoStyle = window.getComputedStyle(ballTwoElement); 
const ball2Animation = __helpers.removeWhiteSpace(ballTwoStyle?.animationTimingFunction);
assert.isTrue(
  ball2Animation == 'ease-out' || ball2Animation == 'cubic-bezier(0,0,0.58,1)'
);
```

# --seed--

## --seed-contents--

```html
<style>

  .balls{
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left: 27%;
    animation-timing-function: linear;
  }
  #ball2 {
    left: 56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }

</style>

<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```

# --solutions--

```html
<style>

  .balls{
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left: 27%;
    animation-timing-function: cubic-bezier(0.25, 0.25, 0.75, 0.75);
  }
  #ball2 {
    left: 56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```
