---
id: 587d78a7367417b2b2512ae0
title: 使用 CSS 動畫來改變按鈕的暫留狀態
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4vZAa'
forumTopicId: 301073
dashedName: use-css-animation-to-change-the-hover-state-of-a-button
---

# --description--

你可以使用 CSS `@keyframes` 來改變按鈕在暫留狀態下的顏色。

以下是暫留時更改影像寬度的範例：

```html
<style>
  img {
    width: 30px;
  }
  img:hover {
    animation-name: width;
    animation-duration: 500ms;
  }

  @keyframes width {
    100% {
      width: 40px;
    }
  }
</style>

<img src="https://cdn.freecodecamp.org/curriculum/applied-visual-design/google-logo.png" alt="Google's Logo" />
```

# --instructions--

請注意，`ms` 代表毫秒，其中 1000ms 相等於 1s。

使用 CSS `@keyframes` 來改變 `button` 元素的 `background-color`，使其在使用者暫留時變為 `#4791d0`。`@keyframes` 規則應該只包含 `100%` 的條目。

# --hints--

@keyframes 規則應該使用 `animation-name` 背景顏色。

```js
assert.match(code, /@keyframes\s+?background-color\s*?{/g);
```

在 `@keyframes` 下應該有一條規則，在 100% 時將 `background-color` 變更為 `#4791d0`。

```js
assert.match(code, /100%\s*?{\s*?background-color:\s*?#4791d0;\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }

  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
  }


</style>

<button>Register</button>
```

# --solutions--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }

  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
  }

  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```
