---
id: 587d78a4367417b2b2512ad2
title: 了解三次色
challengeType: 0
forumTopicId: 301057
dashedName: learn-about-tertiary-colors
---

# --description--

電腦顯示器和裝置螢幕透過組合紅色、綠色和藍色光的量來產生不同的顏色。這在現代色彩理論中稱為 RGB 加色模型。紅色（R）、綠色（G）和藍色（B）稱為原色。混合兩種原色會產生次級色青色（G + B）、洋紅色（R + B）和黃色（R + G）。你在互補色挑戰中看過這些顏色。這些次級色恰好是未用於其產生的原色的互補色，並且在色輪上與該原色相對。例如，洋紅色是由紅色和藍色組成，且是綠色的互補色。

第三色是將一個原色與其相鄰的次色之一混合而成的結果。例如，在 `RGB` 色彩模型中，紅色（原色）和黃色（次色）組合成橙色（第三色）。這為簡單的色輪增加了六種顏色，總共十二種。

有各種選擇不同顏色的方法，可以在設計中產生和諧的組合。一個可以使用第三色的例子稱為分裂互補色結構綱目。此結構綱目從一個基色開始，然後將它與其互補色相鄰的兩個顏色配對。這三種顏色在設計中提供強烈的視覺對比，但比使用兩種互補色更為細膩。

這裡有三種使用分裂互補結構綱目創建的顏色：

<table><thead><tr><th>Color</th><th>Hex Code</th></tr></thead><thead></thead><tbody><tr><td>orange</td><td>#FF7F00</td></tr><tr><td>cyan</td><td>#00FFFF</td></tr><tr><td>raspberry</td><td>#FF007F</td></tr></tbody></table>

# --instructions--

將 `orange`、`cyan` 和 `raspberry` 類別的 `background-color` 屬性更改為它們各自的顏色。請確保使用十六進位色碼，而非顏色名稱。

# --hints--

具有類別 `orange` 的 `div` 元素應該有橘色的 `background-color`。

```js
const orangeElement = document.querySelector('.orange');
const orangeStyle = window.getComputedStyle(orangeElement);
assert.equal(orangeStyle?.backgroundColor, 'rgb(255, 127, 0)');
```

具有類別 `cyan` 的 `div` 元素應該有青色的 `background-color`。

```js
const cyanElement = document.querySelector('.cyan');
const cyanStyle = window.getComputedStyle(cyanElement);
assert.equal(cyanStyle?.backgroundColor, 'rgb(0, 255, 255)');
```

具有類別 `raspberry` 的 `div` 元素應該有覆盆子色的 `background-color`。

```js
const raspberryElement = document.querySelector('.raspberry');
const raspberryStyle = window.getComputedStyle(raspberryElement);
assert.equal(raspberryStyle?.backgroundColor, 'rgb(255, 0, 127)');
```

所有顏色類別的 `background-color` 值應該是十六進位色碼，而非顏色名稱。

```js
assert.notMatch(code, /background-color:\s(orange|cyan|raspberry)/);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .orange {
    background-color: #000000;
  }

  .cyan {
    background-color: #000000;
  }

  .raspberry {
    background-color: #000000;
  }

  div {
    height: 100px;
    width: 100px;
    margin-bottom: 5px;
  }
</style>

<div class="orange"></div>
<div class="cyan"></div>
<div class="raspberry"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .orange {
    background-color: #FF7F00;
  }

  .cyan {
    background-color: #00FFFF;
  }

  .raspberry {
    background-color: #FF007F;
  }

  div {
    height: 100px;
    width: 100px;
    margin-bottom: 5px;
  }
</style>
<div class="orange"></div>
<div class="cyan"></div>
<div class="raspberry"></div>
```
