---
id: 587d78a4367417b2b2512ad5
title: 調整顏色的色調
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDJvT7'
forumTopicId: 301038
dashedName: adjust-the-tone-of-a-color
---

# --description--

CSS 中的 `hsl()` 選項也讓調整顏色的色調變得簡單。將白色與純色調混合會產生該顏色的色調，加入黑色則會形成陰影。或者，透過加入灰色或同時進行調色和陰影，會產生色調。請記得 `hsl()` 中的 's' 和 'l' 分別代表飽和度和亮度。飽和度百分比改變灰色的量，亮度百分比決定顏色中白色或黑色的多少。當你有喜歡的基礎色調，但需要不同變化時，這非常有用。

# --instructions--

所有元素的 `background-color` 預設值為 `transparent`。我們的 `nav` 元素目前看起來有 `cyan` 背景，因為它背後的元素設定了 `background-color` 為 `cyan`。為你的 `nav` 元素添加 `background-color`，使其使用相同的 `cyan` 色調，但具有 `80%` 飽和度和 `25%` 亮度值，以改變其色調和陰影。

# --hints--

`nav` 元素應該使用 `hsl()` 屬性設定調整後青色調的 `background-color`。

```js
// Computed style of hsl(180, 80%, 25%) results in rgb(13,115,115)
assert.equal(
  new __helpers.CSSHelp(document).getStyle('nav').getPropVal('background-color', true), 
  'rgb(13,115,115)'
)
```

# --seed--

## --seed-contents--

```html
<style>
  header {
    background-color: hsl(180, 90%, 35%);
    color: #FFFFFF;
  }

  nav {

  }

  h1 {
    text-indent: 10px;
    padding-top: 10px;
  }

  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }

  nav li {
    display: inline;
    margin-right: 20px;
  }

  a {
    text-decoration: none;
    color: inherit;
  }
</style>

<header>
  <h1>Cooking with FCC!</h1>
  <nav>
    <ul>
      <li><a href="#">Home</a></li>
      <li><a href="#">Classes</a></li>
      <li><a href="#">Contact</a></li>
    </ul>
  </nav>
</header>
```

# --solutions--

```html
<style>
  header {
    background-color: hsl(180, 90%, 35%);
    color: #FFFFFF;
  }

  nav {
    background-color: hsl(180, 80%, 25%);
  }

  h1 {
    text-indent: 10px;
    padding-top: 10px;
  }

  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }

  nav li {
    display: inline;
    margin-right: 20px;
  }

  a {
    text-decoration: none;
    color: inherit;
  }
</style>
<header>
  <h1>Cooking with FCC!</h1>
  <nav>
    <ul>
      <li><a href="#">Home</a></li>
      <li><a href="#">Classes</a></li>
      <li><a href="#">Contact</a></li>
    </ul>
  </nav>
</header>
```
