---
id: 587d781e367417b2b2512acb
title: 使用絕對定位將元素鎖定到其父元素
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLJ7c3'
forumTopicId: 301060
dashedName: lock-an-element-to-its-parent-with-absolute-positioning
---

# --description--

下一個 CSS `position` 屬性的選項是 `absolute`，它會將元素相對於其父容器鎖定在固定位置。與 `relative` 位置不同，這會將元素從文件的正常流程中移除，因此周圍的項目會忽略它。CSS 偏移屬性（top 或 bottom 以及 left 或 right）用於調整位置。

絕對定位的一個細微差別是它會相對於其最近的*已定位*祖先被鎖定。如果你忘記為父項目添加定位規則（這通常是使用 `position: relative;` 完成的），瀏覽器會繼續向上尋找串鏈，最終預設為 `body` 頁籤。

# --instructions--

將 `#searchbar` 元素鎖定在其 `section` 父元素的右上方，方法是將其 `position` 宣告為 `absolute`。為它設定 `top` 和 `right` 偏移量各 50 像素。

# --hints--

`#searchbar` 元素的 `position` 應設為 `absolute`。

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.position, 'absolute');
```

你的程式碼應該在 `#searchbar` 元素上使用 `top` CSS 偏移量 50 像素。

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.top, '50px');
```

你的程式碼應該在 `#searchbar` 元素上使用 `right` CSS 偏移量 50 像素。

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.right, '50px');
```

# --seed--

## --seed-contents--

```html
<style>
  #searchbar {



  }
  section {
    position: relative;
  }
</style>
<body>
  <h1>Welcome!</h1>
  <section>
    <form id="searchbar">
      <label for="search">Search:</label>
      <input type="search" id="search" name="search">
      <input type="submit" name="submit" value="Go!">
    </form>
  </section>
</body>
```

# --solutions--

```html
<style>
  #searchbar {
    position: absolute;
    top: 50px;
    right: 50px;
  }
  section {
    position: relative;
  }
</style>
<body>
  <h1>Welcome!</h1>
  <section>
    <form id="searchbar">
      <label for="search">Search:</label>
      <input type="search" id="search" name="search">
      <input type="submit" name="submit" value="Go!">
    </form>
  </section>
</body>
```
