---
id: 587d78a5367417b2b2512ad6
title: 創建漸進式 CSS 線性漸層
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4dpt9'
forumTopicId: 301047
dashedName: create-a-gradual-css-linear-gradient
---

# --description--

在 HTML 元素上套用顏色不限於單一純色。CSS 提供了在元素上使用顏色漸層（又稱為漸變）的能力。這是透過 `background` 屬性的 `linear-gradient()` 函式（程式）來存取的。以下是一般語法：

```css
background: linear-gradient(gradient_direction, color 1, color 2, color 3, ...);
```

第一個引數指定顏色漸層開始的方向－它可以以角度表示，其中 `90deg` 會產生水平漸層（從左到右），而 `45deg` 會產生對角漸層（從左下到右上）。接下來的引數指定漸層中使用的顏色順序。

範例：

```css
background: linear-gradient(90deg, red, yellow, rgb(204, 204, 255));
```

# --instructions--

為你的 `div` 元素的 `background` 使用 `linear-gradient()`，並將方向設為 35 度，使顏色從 `#CCFFFF` 變更為 `#FFCCCC`。

# --hints--

`div` 元素應該有指定方向和顏色的 `linear-gradient` `background`。

```js
const divElement = document.querySelector('div');
const divStyle = window.getComputedStyle(divElement); 
assert.match(divStyle?.background, /linear-gradient\(35deg, rgb\(204, 255, 255\), rgb\(255, 204, 204\)\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin: 50px auto;

  }

</style>

<div></div>
```

# --solutions--

```html
<style>
  div {
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin: 50px auto;
    background: linear-gradient(35deg, #CCFFFF, #FFCCCC);
  }
</style>
<div></div>
```
