---
id: 587d78a6367417b2b2512adb
title: 使用 CSS Transform 屬性 skewX 沿 X 軸傾斜元素
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLP8Sr'
forumTopicId: 301074
dashedName: use-the-css-transform-property-skewx-to-skew-an-element-along-the-x-axis
---

# --description--

`transform` 屬性的下一個函式（程式）是 `skewX()`，它會將選取的元素沿其 X（水平）軸傾斜指定的角度。

以下 `code` 將段落元素沿 X 軸傾斜 -32 度。

```css
p {
  transform: skewX(-32deg);
}
```

# --instructions--

使用 `transform` 屬性將 id 為 `bottom` 的元素沿 X 軸傾斜 24 度。

# --hints--

id 為 `bottom` 的元素應沿其 X 軸傾斜 24 度。

```js
assert.match(code, /#bottom\s*?{\s*?.*?\s*?transform:\s*?skewX\(24deg\);/g);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
  }
  #top {
    background-color: red;
  }
  #bottom {
    background-color: blue;

  }
</style>

<div id="top"></div>
<div id="bottom"></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
  }
  #top {
    background-color: red;
  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>
<div id="top"></div>
<div id="bottom"></div>
```
