---
id: 587d78a3367417b2b2512ace
title: 使用 float 屬性將元素向左或向右推動
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MDqu2'
forumTopicId: 301066
dashedName: push-elements-left-or-right-with-the-float-property
---

# --description--

下一個定位工具實際上不使用 `position`，而是設定元素的 `float` 屬性。浮動元素會從文件的正常流程中移除，並推向其包含父元素的 `left` 或 `right`。通常會搭配 `width` 屬性使用，以指定浮動元素所需的水平空間。

# --instructions--

給定的標記語言非常適合作為兩行版面配置，`section` 和 `aside` 元素並排顯示。為你的 `#left` 項目添加 `float` 的 `left`，為你的 `#right` 項目添加 `float` 的 `right`。

# --hints--

id 為 `left` 的元素應該有 `float` 值為 `left`。

```js
const leftElement = document.querySelector('#left');
const leftStyle = window.getComputedStyle(leftElement);
assert.equal(leftStyle?.float, 'left');
```

id 為 `right` 的元素應該有 `float` 值為 `right`。

```js
const rightElement = document.querySelector('#right');
const rightStyle = window.getComputedStyle(rightElement);
assert.equal(rightStyle?.float, 'right');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
    #left {

      width: 50%;
    }
    #right {

      width: 40%;
    }
    aside, section {
      padding: 2px;
      background-color: #ccc;
    }
  </style>
</head>
<body>
  <header>
    <h1>Welcome!</h1>
  </header>
  <section id="left">
    <h2>Content</h2>
    <p>Good stuff</p>
  </section>
  <aside id="right">
    <h2>Sidebar</h2>
    <p>Links</p>
  </aside>
</body>
```

# --solutions--

```html
<head>
  <style>
    #left {
      float: left;
      width: 50%;
    }
    #right {
      float: right;
      width: 40%;
    }
    aside, section {
      padding: 2px;
      background-color: #ccc;
    }
  </style>
</head>
<body>
  <header>
    <h1>Welcome!</h1>
  </header>
  <section id="left">
    <h2>Content</h2>
    <p>Good stuff</p>
  </section>
  <aside id="right">
    <h2>Sidebar</h2>
    <p>Links</p>
  </aside>
</body>
```
