---
id: 587d78a4367417b2b2512ad4
title: 調整顏色的色相
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp38TZ'
forumTopicId: 301036
dashedName: adjust-the-hue-of-a-color
---

# --description--

顏色有多個特性，包括色相、飽和度和亮度。CSS3 介紹了 `hsl()` 函式，作為直接說明這些特性來選擇顏色的替代方法。

**色相** 是人們通常所認為的「顏色」。如果你想像一個色譜，左邊是紅色，中間是綠色，右邊是藍色，色相就是顏色在這條線上的位置。在 `hsl()` 中，色相使用色輪的概念，而非色譜，顏色在圓上的角度以 0 到 360 之間的數值表示。

**飽和度** 是顏色中灰色的量。完全飽和的顏色沒有灰色，而最低飽和的顏色幾乎完全是灰色。這以百分比表示，100％ 表示完全飽和。

**明度** 是顏色中白色或黑色的量。百分比範圍從 0%（黑色）到 100%（白色），其中 50% 是正常顏色。

以下是使用完全飽和、正常亮度色彩的 `hsl()` 範例：

<table><thead><tr><th>Color</th><th>HSL</th></tr></thead><tbody><tr><td>red</td><td>hsl(0, 100%, 50%)</td></tr><tr><td>yellow</td><td>hsl(60, 100%, 50%)</td></tr><tr><td>green</td><td>hsl(120, 100%, 50%)</td></tr><tr><td>cyan</td><td>hsl(180, 100%, 50%)</td></tr><tr><td>blue</td><td>hsl(240, 100%, 50%)</td></tr><tr><td>magenta</td><td>hsl(300, 100%, 50%)</td></tr></tbody></table>

# --instructions--

根據類別名稱（`green`、`cyan` 或 `blue`）使用 `hsl()` 變更每個 `div` 元素的 `background-color`。三者皆應具有完全飽和度和正常亮度。

# --hints--

你的程式碼應該使用 `hsl()` 函式來宣告綠色。

```js
assert.match(code,/\.green\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

你的程式碼應該使用 `hsl()` 函式來宣告青色。

```js
assert.match(code,/\.cyan\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

你的程式碼應該使用 `hsl()` 函式來宣告藍色。

```js
assert.match(code,/\.blue\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

具有類別 `green` 的 `div` 元素應該有綠色的 `background-color`。

```js
const greenElement = document.querySelector(".green");
const greenStyle = window.getComputedStyle(greenElement); 
assert.equal(greenStyle?.backgroundColor, 'rgb(0, 255, 0)');
```

具有類別 `cyan` 的 `div` 元素應該有青色的 `background-color`。

```js
const cyanElement = document.querySelector(".cyan");
const cyanStyle = window.getComputedStyle(cyanElement); 
assert.equal(cyanStyle?.backgroundColor, 'rgb(0, 255, 255)');
```

具有類別 `blue` 的 `div` 元素應該有藍色的 `background-color`。

```js
const blueElement = document.querySelector(".blue");
const blueStyle = window.getComputedStyle(blueElement); 
assert.equal(blueStyle?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .green {
    background-color: #000000;
  }

  .cyan {
    background-color: #000000;
  }

  .blue {
    background-color: #000000;
  }

  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>

<div class="green"></div>
<div class="cyan"></div>
<div class="blue"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .green {
    background-color: hsl(120, 100%, 50%);
  }

  .cyan {
    background-color:   hsl(180, 100%, 50%);
  }

  .blue {
    background-color: hsl(240, 100%, 50%);
  }

  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>
<div class="green"></div>
<div class="cyan"></div>
<div class="blue"></div>
```
