---
id: 587d78a7367417b2b2512ae1
title: 使用 CSS 動畫創建移動效果
challengeType: 0
videoUrl: 'https://scrimba.com/c/c7amZfW'
forumTopicId: 301051
dashedName: create-movement-using-css-animation
---

# --description--

當元素具有指定的 `position`，例如 `fixed` 或 `relative`，可以在動畫規則中使用 CSS 偏移屬性 `right`、`left`、`top` 和 `bottom` 來創建移動效果。

如以下範例所示，你可以透過將 `50%` 訊框的 `top` 屬性設定為 50px，然後將第一個（`0%`）和最後一個（`100%`）訊框的 `top` 屬性設定為 0px，來將該項目先向下推再向上推。

```css
@keyframes rainbow {
  0% {
    background-color: blue;
    top: 0px;
  }
  50% {
    background-color: green;
    top: 50px;
  }
  100% {
    background-color: yellow;
    top: 0px;
  }
}
```

# --instructions--

為你的 `div` 動畫添加水平移動。使用 `left` 偏移屬性，將其添加到 `@keyframes` 規則中，使 rainbow 在 `0%` 時從 0 像素開始，於 `50%` 時移動到 25 像素，並在 `100%` 時結束於 -25 像素。不要在編輯器中取代 `top` 屬性－動畫應該同時具有垂直和水平移動。

# --hints--

`0%` 的 `@keyframes` 規則應該使用 `left` 偏移量 0px。

```js
assert.match(code,/[^50]0%\s*?{[\s\S]*?left:\s*?0px(;[\s\S]*?|\s*?)}/gi);
```

`@keyframes` 規則中 `50%` 應該使用 25px 的 `left` 偏移量。

```js
assert.match(code,/50%\s*?{[\s\S]*?left:\s*?25px(;[\s\S]*?|\s*?)}/gi);
```

`@keyframes` 規則中 `100%` 應該使用 `left` 偏移量 -25px。

```js
assert.match(code,/100%\s*?{[\s\S]*?left:\s*?-25px(;[\s\S]*?|\s*?)}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
    position: relative;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
      top: 0px;

    }
    50% {
      background-color: green;
      top: 50px;

    }
    100% {
      background-color: yellow;
      top: 0px;

    }
  }
</style>

<div id="rect"></div>
```

# --solutions--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
    position: relative;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
      top: 0px;
      left: 0px;
    }
    50% {
      background-color: green;
      top: 50px;
      left: 25px;
    }
    100% {
      background-color: yellow;
      top: 0px;
      left: -25px;
    }
  }
</style>
<div id="rect"></div>
```
