---
id: 587d781e367417b2b2512aca
title: 使用 CSS 偏移量移動相對定位的元素
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bQEA4'
forumTopicId: 301065
dashedName: move-a-relatively-positioned-element-with-css-offsets
---

# --description--

`top` 或 `bottom`，以及 `left` 或 `right` 的 CSS 偏移量告訴瀏覽器相對於該項目在文件正常流程中所處位置，要偏移多遠。你正在將元素從指定位置偏移，這會使元素遠離被引用的邊（實際上是相反方向）。如你在上一個挑戰中看到的，使用 `top` 偏移會使 `h2` 向下移動。同樣地，使用 `left` 偏移會使項目向右移動。

# --instructions--

使用 CSS 偏移將 `h2` 向右移動 15 像素，向上移動 10 像素。

# --hints--

你的程式碼應該使用 CSS 偏移量，將 `h2` 相對定位向上移動 10 px。換句話說，將它從正常位置的 `bottom` 向上移動 10 px。

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.bottom, '10px');
```

你的程式碼應該使用 CSS 偏移量，將 `h2` 相對定位向右移動 15px。換句話說，將它從原本所在位置的 `left` 向外移動 15px。

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.left , '15px');
```

# --seed--

## --seed-contents--

```html
<head>
<style>
  h2 {
    position: relative;


  }
</style>
</head>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```

# --solutions--

```html
<head>
<style>
  h2 {
    position: relative;
    left: 15px;
    bottom: 10px;
  }
</style>
</head>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```
