---
id: 587d78a3367417b2b2512acf
title: 使用 z-index 屬性變更重疊元素的位置
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM94aHk'
forumTopicId: 301046
dashedName: change-the-position-of-overlapping-elements-with-the-z-index-property
---

# --description--

當元素被定位為重疊（例如使用 `position: absolute | relative | fixed | sticky`）時，預設情況下，HTML 標記中後出現的元素會顯示在其他元素之上。然而，`z-index` 屬性可以指定元素堆疊的順序。它必須是整數（即整數而非小數），且元素的 `z-index` 屬性值越高，該元素在堆疊中的位置就越高，超過值較低的元素。

# --instructions--

為你的類別名稱為 `first`（紅色矩形）的元素添加 `z-index` 屬性，並將其設為 2，以便它涵蓋其他元素（藍色矩形）。

# --hints--

具有類別 `first` 的元素應該有 `z-index` 值為 2。

```js
const firstElement = document.querySelector('.first');
const firstStyle = window.getComputedStyle(firstElement); 
assert.equal(firstStyle?.zIndex, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 60%;
    height: 200px;
    margin-top: 20px;
  }

  .first {
    background-color: red;
    position: absolute;

  }
  .second {
    background-color: blue;
    position: absolute;
    left: 40px;
    top: 50px;
    z-index: 1;
  }
</style>

<div class="first"></div>
<div class="second"></div>
```

# --solutions--

```html
<style>
  div {
    width: 60%;
    height: 200px;
    margin-top: 20px;
  }

  .first {
    background-color: red;
    position: absolute;
    z-index: 2;
  }
  .second {
    background-color: blue;
    position: absolute;
    left: 40px;
    top: 50px;
    z-index: 1;
  }
</style>
<div class="first"></div>
<div class="second"></div>
```
