---
id: 587d78a6367417b2b2512ade
title: 使用 CSS 和 HTML 創建更複雜的形狀
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPpz4fr'
forumTopicId: 301050
dashedName: create-a-more-complex-shape-using-css-and-html
---

# --description--

世界上最受歡迎的形狀之一是愛心形狀，在這個挑戰中你將使用純 CSS 來創建一個愛心。但首先，你需要了解 `::before` 和 `::after` 偽元素。`::before` 會創建一個偽元素，該元素是所選元素的第一個子元素；`::after` 會創建一個偽元素，該元素是所選元素的最後一個子元素。在以下範例中，使用 `::before` 偽元素為具有 `heart` 類別的元素添加一個矩形：

```css
.heart::before {
  content: "";
  background-color: yellow;
  border-radius: 25%;
  position: absolute;
  height: 50px;
  width: 70px;
  top: -50px;
  left: 5px;
}
```

為了讓 `::before` 和 `::after` 偽元素正常運作，它們必須有定義的 `content` 屬性。這個屬性通常用來為所選元素添加像是照片或文字的東西。當使用 `::before` 和 `::after` 偽元素來製作形狀時，仍然需要 `content` 屬性，但它被設為空字串。在上述範例中，具有 `heart` 類別的元素有一個 `::before` 偽元素，該偽元素產生一個高度和寬度分別為 `50px` 和 `70px` 的黃色矩形。這個矩形因為 25% 的 `border-radius` 而有圓角，並且絕對定位在距離元素左側 `5px` 和上方 `50px` 的位置。

# --instructions--

將螢幕上的元素轉換成愛心。在 `.heart::after` 選擇器中，將 `background-color` 改為 `pink`，並將 `border-radius` 改為 50%。

接著，標的為類別為 `heart`（僅 `heart`）的元素，並填入 `transform` 屬性。使用角度為 -45 度的 `rotate()` 函式（程式）。

最後，在 `.heart::before` 選擇器中，將其 `content` 屬性設為空字串。

# --hints--

`.heart::after` 選擇器的 `background-color` 屬性應該是 `pink`。

```js
const heartAfter = code.match(/\.heart::after\s*{[\s\S]+?[^\}]}/g)?.[0];
assert.match(heartAfter, /({|;)\s*background-color\s*:\s*pink\s*(;|})/g);
```

`.heart::after` 選擇器的 `border-radius` 應該是 50%。

```js
assert.lengthOf(code.match(/border-radius\s*?:\s*?50%/gi),2);
```

`heart` 類別的 `transform` 屬性應該使用設定為 -45 度的 `rotate()` 函式（程式）。

```js
assert.match(code,/transform\s*?:\s*?rotate\(\s*?-45deg\s*?\)/gi);
```

`.heart::before` 選擇器的 `content` 應該是空字串。

```js
assert.match(code,/\.heart::before\s*?{\s*?content\s*?:\s*?("|')\1\s*?;/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: ;
  }
  .heart::after {
    background-color: blue;
    content: "";
    border-radius: 25%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart::before {
    content: ;
    background-color: pink;
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }
</style>
<div class="heart"></div>
```

# --solutions--

```html
<style>
  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: rotate(-45deg);
  }
  .heart::after {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart::before {
    content: "";
    background-color: pink;
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }
</style>
<div class="heart"></div>
```
