---
id: 587d78a8367417b2b2512ae7
title: 使用關鍵字變更動畫時間
challengeType: 0
videoUrl: 'https://scrimba.com/c/cJKvwCM'
forumTopicId: 301045
dashedName: change-animation-timing-with-keywords
---

# --description--

在 CSS 動畫中，`animation-timing-function` 屬性控制動畫元素在動畫持續時間內變化的速度。如果動畫是一輛車在給定時間（你的 `animation-duration`）內從點 A 移動到點 B，`animation-timing-function` 就決定這輛車在整個行駛過程中的加速和減速方式。

有許多為熱門選項預先定義的關鍵字。 例如，預設值是 `ease`，它開始時緩慢，中間加速，最後再減速。 其他選項包括 `ease-out`，開始時快速然後減速，`ease-in`，開始時緩慢，最後加速，或 `linear`，整個動畫過程中保持恆定速度。

# --instructions--

對於 id 為 `ball1` 和 `ball2` 的元素，為每個元素添加一個 `animation-timing-function` 屬性，並將 `#ball1` 設定為 `linear`，`#ball2` 設定為 `ease-out`。注意元素在動畫過程中的移動方式不同，但因為它們共用相同的 2 秒 `animation-duration`，所以最終會一起結束。

# --hints--

id 為 `ball1` 的元素的 `animation-timing-function` 屬性值應該是 `linear`。

```js
const ballOne =document.querySelector('#ball1'); 
const ballOneStyle = window.getComputedStyle(ballOne); 

const ball1Animation = __helpers.removeWhiteSpace(ballOneStyle?.animationTimingFunction);
assert.isTrue(ball1Animation == 'linear' || ball1Animation == 'cubic-bezier(0,0,1,1)');
```

id 為 `ball2` 的元素的 `animation-timing-function` 屬性值應該是 `ease-out`。

```js
const ballTwo = document.querySelector('#ball2'); 
const ballTwoStyle = window.getComputedStyle(ballTwo); 

const ball2Animation = __helpers.removeWhiteSpace(ballTwoStyle?.animationTimingFunction);
assert.isTrue(ball2Animation == 'ease-out' || ball2Animation == 'cubic-bezier(0,0,0.58,1)');
```

# --seed--

## --seed-contents--

```html
<style>

  .balls {
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left:27%;

  }
  #ball2 {
    left:56%;

  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }

</style>

<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```

# --solutions--

```html
<style>
  .balls {
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left:27%;
    animation-timing-function: linear;
  }
  #ball2 {
    left:56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```
