---
id: 587d78a9367417b2b2512ae9
title: 使用貝茲曲線移動圖形
challengeType: 0
videoUrl: 'https://scrimba.com/c/c6bnRCK'
forumTopicId: 301071
dashedName: use-a-bezier-curve-to-move-a-graphic
---

# --description--

先前的挑戰討論了描述動畫變化先加速後在動畫結束時減速的 `ease-out` 關鍵字。在右側，展示了 `ease-out` 關鍵字（用於藍色元素）與 `linear` 關鍵字（用於紅色元素）之間的差異。使用訂製的三次貝茲爾曲線函式（function）可以達成與 `ease-out` 關鍵字類似的動畫進展。

一般來說，改變 `p1` 和 `p2` 錨點會驅動不同貝茲曲線的產生，這控制動畫如何隨時間推進。以下是一個使用數值模擬 `ease-out` 風格的貝茲曲線範例：

```css
animation-timing-function: cubic-bezier(0, 0, 0.58, 1);
```

請記住，所有的 `cubic-bezier` 函式（程式）都從 (0, 0) 的 `p0` 開始，並以 (1, 1) 的 `p3` 結束。在此範例中，曲線在 Y 軸上的移動速度較快（從 0 開始，移動到 `p1` 的 y 值為 0，然後移動到 `p2` 的 y 值為 1），而在 X 軸上的移動速度較慢（從 0 開始，然後 `p1` 為 0，`p2` 為 0.58）。因此，動畫元素的變化進度比該段動畫的時間更快。在曲線接近尾端時，x 與 y 值變化的關係反轉——y 值從 1 移動到 1（無變化），而 x 值從 0.58 移動到 1，使動畫變化的進度相較於動畫持續時間變慢。

# --instructions--

要看到這個貝茲曲線的實際效果，將 id 為 `red` 的元素的 `animation-timing-function` 更改為一個 `cubic-bezier` 函式，並將 x1、y1、x2、y2 分別設為 `0, 0, 0.58, 1` 作為值。這將使兩個元素以類似的方式進行動畫。

# --hints--

id 為 `red` 的元素的 `animation-timing-function` 屬性值應該是一個 `cubic-bezier` 函式，其 x1、y1、x2、y2 值分別設定為 0、0、0.58、1。

```js
const redElement = document.querySelector('#red');
const redStyle = window.getComputedStyle(redElement);
assert.equal(
  redStyle?.animationTimingFunction, 'cubic-bezier(0, 0, 0.58, 1)'
);
```

id 為 `red` 的元素不應再具有 `animation-timing-function` 屬性為 `linear`。

```js
const redElement = document.querySelector('#red');
const redStyle = window.getComputedStyle(redElement);
assert.notEqual(redStyle?.animationTimingFunction, 'linear');
```

id 為 `blue` 的元素的 `animation-timing-function` 屬性值不應該改變。

```js
const blueElement = document.querySelector('#blue');
const blueStyle = window.getComputedStyle( blueElement);
const blueBallAnimation = __helpers.removeWhiteSpace(
  blueStyle?.animationTimingFunction
);
assert.isTrue(
  blueBallAnimation == 'ease-out' ||
    blueBallAnimation == 'cubic-bezier(0,0,0.58,1)'
);
```

# --seed--

## --seed-contents--

```html
<style>
  .balls{
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 27%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 56%;
    animation-timing-function: ease-out;
  }
  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id= "red"></div>
<div class="balls" id= "blue"></div>
```

# --solutions--

```html
<style>
  .balls{
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 27%;
    animation-timing-function: cubic-bezier(0, 0, 0.58, 1);
  }
  #blue {
    background: blue;
    left: 56%;
    animation-timing-function: ease-out;
  }
  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id= "red"></div>
<div class="balls" id= "blue"></div>
```
