---
id: 587d78a7367417b2b2512ae2
title: 透過從左到右淡出元素來創建視覺導向
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJqqAE'
forumTopicId: 301054
dashedName: create-visual-direction-by-fading-an-element-from-left-to-right
---

# --description--

在這個挑戰中，你將改變動畫元素的 `opacity`，使其在到達螢幕右側時逐漸淡出。

在顯示的動畫中，帶有漸層背景的圓形元素根據 `@keyframes` 規則向右移動至動畫的 50% 標記。

# --instructions--

以 id 為 `ball` 的元素為標的，並在 `50%` 時將 `opacity` 屬性設為 0.1，使該元素在向右移動時逐漸淡出。

# --hints--

`keyframes` 規則中，fade 應該在 50% 時將 `opacity` 屬性設為 0.1。

```js
assert.match(code,/@keyframes fade\s*?{\s*?50%\s*?{\s*?(?:left:\s*?60%;\s*?opacity:\s*?0?\.1;|opacity:\s*?0?\.1;\s*?left:\s*?60%;)/gi
);
```

# --seed--

## --seed-contents--

```html
<style>

  #ball {
    width: 70px;
    height: 70px;
    margin: 50px auto;
    position: fixed;
    left: 20%;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: fade;
    animation-duration: 3s;
  }

  @keyframes fade {
    50% {
      left: 60%;

    }
  }

</style>

<div id="ball"></div>
```

# --solutions--

```html
<style>
  #ball {
    width: 70px;
    height: 70px;
    margin: 50px auto;
    position: fixed;
    left: 20%;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: fade;
    animation-duration: 3s;
  }

  @keyframes fade {
    50% {
      left: 60%;
      opacity: 0.1;
    }
  }
</style>
<div id="ball"></div>
```
