---
id: 587d78a6367417b2b2512adc
title: 使用 CSS Transform 屬性 skewY 沿 Y 軸傾斜元素
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MZ2uB'
forumTopicId: 301075
dashedName: use-the-css-transform-property-skewy-to-skew-an-element-along-the-y-axis
---

# --description--

鑑於 `skewX()` 函式（程式）會將所選元素沿 X 軸傾斜指定角度，因此 `skewY()` 屬性會將元素沿 Y（垂直）軸傾斜也不足為奇。

# --instructions--

使用 `transform` 屬性將 id 為 `top` 的元素沿 Y 軸傾斜 -10 度。

# --hints--

id 為 `top` 的元素應沿其 Y 軸傾斜 -10 度。

```js
assert.match(code, /#top\s*?{\s*?.*?\s*?transform:\s*?skewY\(-10deg\);/g);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin: 50px auto;
  }
  #top {
    background-color: red;

  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>

<div id="top"></div>
<div id="bottom"></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin: 50px auto;
  }
  #top {
    background-color: red;
    transform: skewY(-10deg);
  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>
<div id="top"></div>
<div id="bottom"></div>
```
