---
id: 587d78a5367417b2b2512ad8
title: 透過將細微樣式作為背景影像來創建材質
challengeType: 0
videoUrl: 'https://scrimba.com/c/cQdwJC8'
forumTopicId: 301052
dashedName: create-texture-by-adding-a-subtle-pattern-as-a-background-image
---

# --description--

為背景添加質感和趣味，並讓它更突出的一種方法是添加細微的樣式。關鍵在於平衡，因為你不希望背景過於突出，搶走前景的焦點。`background` 屬性支援 `url()` 函式（程式），用以聯結到所選質感或樣式的影像。聯結位址包裹在小括弧內的引號中。

# --instructions--

使用 `https://cdn-media-1.freecodecamp.org/imgr/MJAkxbh.png` 的 url，使用 `body` 選擇器設定整個頁面的 `background`。

# --hints--

你的 `body` 元素應該有一個 `background` 屬性，設定為帶有指定聯結的 `url()`。

```js
assert.match(code,/background(-image)?:\s*?url\(\s*("|'|)https:\/\/cdn-media-1\.freecodecamp\.org\/imgr\/MJAkxbh\.png\2\s*\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {

  }
</style>
```

# --solutions--

```html
<style>
  body {
    background: url("https://cdn-media-1.freecodecamp.org/imgr/MJAkxbh.png");
  }
</style>
```
