---
id: 587d78a8367417b2b2512ae6
title: 以不同速率動畫多個元素
challengeType: 0
videoUrl: 'https://scrimba.com/c/cnpWZc9'
forumTopicId: 301042
dashedName: animate-multiple-elements-at-variable-rates
---

# --description--

在前一個挑戰中，你透過更改兩個類似動畫元素的 `@keyframes` 規則來改變動畫速率。你也可以透過操作多個元素的 `animation-duration` 來達成相同的目標。

在程式碼編輯器中執行的動畫裡，天空中有三顆星星以相同速率持續迴圈閃爍。若要讓它們以不同速率閃爍，你可以為每個元素設定不同的 `animation-duration` 屬性值。

# --instructions--

將類別為 `star-1`、`star-2` 和 `star-3` 的元素的 `animation-duration` 分別設定為 1s、0.9s 和 1.1s。

# --hints--

類別為 `star-1` 的星星的 `animation-duration` 屬性應保持在 1s。

```js
const starOne = document.querySelector('.star-1');
const starOneStyle = window.getComputedStyle(starOne); 
assert.equal(starOneStyle?.animationDuration, '1s');
```

類別為 `star-2` 的星星的 `animation-duration` 屬性應該是 0.9s。

```js
const starTwo = document.querySelector('.star-2');
const starTwoStyle = window.getComputedStyle(starTwo); 
assert.equal(starTwoStyle?.animationDuration , '0.9s');
```

類別為 `star-3` 的星星的 `animation-duration` 屬性應該是 1.1s。

```js
const starThree = document.querySelector('.star-3');
const starThreeStyle = window.getComputedStyle(starThree); 
assert.equal(starThreeStyle?.animationDuration, '1.1s');
```

# --seed--

## --seed-contents--

```html
<style>
  .stars {
    background-color: white;
    height: 30px;
    width: 30px;
    border-radius: 50%;
    animation-iteration-count: infinite;
  }

  .star-1 {
    margin-top: 15%;
    margin-left: 60%;
    animation-duration: 1s;
    animation-name: twinkle;
  }

  .star-2 {
    margin-top: 25%;
    margin-left: 25%;
    animation-duration: 1s;
    animation-name: twinkle;
  }

  .star-3 {
    margin-top: 10%;
    margin-left: 50%;
    animation-duration: 1s;
    animation-name: twinkle;
  }

  @keyframes twinkle {
    20% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  #back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: linear-gradient(black, #000099, #66c2ff, #ffcccc, #ffeee6);
  }
</style>

<div id="back"></div>
<div class="star-1 stars"></div>
<div class="star-2 stars"></div>
<div class="star-3 stars"></div>
```

# --solutions--

```html
<style>
  .stars {
    background-color: white;
    height: 30px;
    width: 30px;
    border-radius: 50%;
    animation-iteration-count: infinite;
  }

  .star-1 {
    margin-top: 15%;
    margin-left: 60%;
    animation-duration: 1s;
    animation-name: twinkle;
  }

  .star-2 {
    margin-top: 25%;
    margin-left: 25%;
    animation-duration: 0.9s;
    animation-name: twinkle;
  }

  .star-3 {
    margin-top: 10%;
    margin-left: 50%;
    animation-duration: 1.1s;
    animation-name: twinkle;
  }

  @keyframes twinkle {
    20% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  #back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: linear-gradient(black, #000099, #66c2ff, #ffcccc, #ffeee6);
  }
</style>
<div id="back"></div>
<div class="star-1 stars"></div>
<div class="star-2 stars"></div>
<div class="star-3 stars"></div>
```
