---
id: 587d781c367417b2b2512ac2
title: 設定多個標題元素的字型大小
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPpQNT3'
forumTopicId: 301067
dashedName: set-the-font-size-for-multiple-heading-elements
---

# --description--

`font-size` 屬性用於指定給定元素中文字的大小。此規則可用於多個元素，以在頁面上創建文字的視覺一致性。在此挑戰中，你將為所有 `h1` 到 `h6` 頁籤設定值，以平衡標題大小。

# --instructions--

在 `style` 頁籤中，設定以下的 `font-size`：

- `h1` 頁籤高度為 68px。
- `h2` 頁籤高度為 52px。
- `h3` 頁籤高度為 40px。
- `h4` 頁籤字型大小為 32px。
- `h5` 頁籤字型大小為 21px。
- `h6` 頁籤字型大小為 14px。

# --hints--

你的程式碼應該將 `h1` 頁籤的 `font-size` 屬性設定為 68 像素。

```js
 const fontSizeOfh1 = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh1 ,'68px');
```

你的程式碼應該將 `h2` 頁籤的 `font-size` 屬性設定為 52 像素。

```js
 const fontSizeOfh2 = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh2 ,'52px');
```

你的程式碼應該將 `h3` 頁籤的 `font-size` 屬性設為 40 像素。

```js
 const fontSizeOfh3 = new __helpers.CSSHelp(document).getStyle('h3')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh3 ,'40px');
```

你的程式碼應該將 `h4` 頁籤的 `font-size` 屬性設定為 32 像素。

```js
 const fontSizeOfh4 = new __helpers.CSSHelp(document).getStyle('h4')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh4 , '32px');
```

你的程式碼應該將 `h5` 頁籤的 `font-size` 屬性設定為 21 像素。

```js
 const fontSizeOfh5 = new __helpers.CSSHelp(document).getStyle('h5')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh5 ,'21px');
```

你的程式碼應該將 `h6` 頁籤的 `font-size` 屬性設為 14 像素。

```js
 const fontSizeOfh6 = new __helpers.CSSHelp(document).getStyle('h6')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh6 , '14px');
```

# --seed--

## --seed-contents--

```html
<style>






</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```

# --solutions--

```html
<style>
  h1 {
    font-size: 68px;
  }
  h2 {
    font-size: 52px;
  }
  h3 {
    font-size: 40px;
  }
  h4 {
    font-size: 32px;
  }
  h5 {
    font-size: 21px;
  }
  h6 {
    font-size: 14px;
  }
</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```
