---
id: 587d781e367417b2b2512acc
title: 使用固定定位將元素鎖定到瀏覽器視窗
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MDNUR'
forumTopicId: 301061
dashedName: lock-an-element-to-the-browser-window-with-fixed-positioning
---

# --description--

CSS 提供的下一種佈局結構綱目是 `fixed` 位置，這是一種絕對定位的型別，會將元素鎖定在瀏覽器視窗的相對位置。與絕對定位類似，它與 CSS 偏移屬性一起使用，並且也會將元素從文件的正常流程中移除。其他項目不再「察覺」它的位置，這可能需要在其他地方進行一些佈局調整。

`fixed` 和 `absolute` 位置之間的一個主要差異是，具有固定位置的元素在使用者滾動時不會移動。

# --instructions--

程式碼中的導覽列標籤為 `navbar` 的 id。將它的 `position` 改為 `fixed`，並將它從 `top` 偏移 0 像素，從 `left` 偏移 0 像素。當你添加完程式碼後，捲動預覽視窗以查看導覽列如何保持固定位置。

# --hints--

`#navbar` 元素的 `position` 應該設為 `fixed`。

```js
const navbarElement = document.querySelector('#navbar');
const navbarStyle = window.getComputedStyle(navbarElement);
 assert.equal(navbarStyle?.position, 'fixed');
```

你的程式碼應該在 `#navbar` 元素上使用 0 像素的 `top` CSS 偏移量。

```js
const navbarElement = document.querySelector('#navbar');
 const navbarStyle = window.getComputedStyle(navbarElement);
assert.equal(navbarStyle?.top, '0px');
```

你的程式碼應該在 `#navbar` 元素上使用 `left` CSS 偏移量為 0 像素。

```js
const navbarElement = document.querySelector('#navbar');
 const navbarStyle = window.getComputedStyle(navbarElement);
 assert.equal(navbarStyle?.left, '0px');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    min-height: 150vh;
  }
  #navbar {



    width: 100%;
    background-color: #767676;
  }
  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }
  nav li {
    display: inline;
    margin-right: 20px;
  }
  a {
    text-decoration: none;
  }
</style>
<body>
  <header>
    <h1>Welcome!</h1>
    <nav id="navbar">
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Contact</a></li>
      </ul>
    </nav>
  </header>
  <p>I shift up when the #navbar is fixed to the browser window.</p>
</body>
```

# --solutions--

```html
<style>
  body {
    min-height: 150vh;
  }
  #navbar {
    position: fixed;
    top: 0;
    left: 0;
    width: 100%;
    background-color: #767676;
  }
  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }
  nav li {
    display: inline;
    margin-right: 20px;
  }
  a {
    text-decoration: none;
  }
</style>
<body>
  <header>
    <h1>Welcome!</h1>
    <nav id="navbar">
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Contact</a></li>
      </ul>
    </nav>
  </header>
  <p>I shift up when the #navbar is fixed to the browser window.</p>
</body>
```
