---
id: 58a7a6ebf9a6318348e2d5aa
title: 修改動畫的 Fill Mode
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJDmcE'
forumTopicId: 301064
dashedName: modify-fill-mode-of-an-animation
---

# --description--

那很好，但目前還無法正常運作。請注意動畫在 `500ms` 過後會重置，導致按鈕回復到原本的顏色。你希望按鈕保持凸顯標示。

這可以透過將 `animation-fill-mode` 屬性設定為 `forwards` 來完成。`animation-fill-mode` 指定動畫結束後套用到元素的樣式。你可以這樣設定：

```css
animation-fill-mode: forwards;
```

# --instructions--

將 `button:hover` 的 `animation-fill-mode` 屬性設定為 `forwards`，以便使用者暫留時按鈕保持凸顯標示。

# --hints--

`button:hover` 應該有一個 `animation-fill-mode` 屬性，值為 `forwards`。

```js
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-fill-mode\s*?:\s*?forwards\s*?;[\s\S]*}/gi);
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-name\s*?:\s*?background-color\s*?;[\s\S]*}/gi);
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-duration\s*?:\s*?500ms\s*?;[\s\S]*}/gi );
```

# --seed--

## --seed-contents--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
    /* Only change code below this line */

    /* Only change code above this line */
  }
  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```

# --solutions--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
    animation-fill-mode: forwards;
  }
  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```
