---
id: 6a3e6c0eb5d050ce07781709
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

`if` 조건이 `False`로 평가되면(즉, 충분한 달걀이 있다는 뜻), 함수는 `if` 블록을 건너뛰고 나머지 코드를 계속 실행합니다.

이제 데이터를 처리하고 새로 계산된 값을 반환해야 합니다.

`if` 블록 바깥에서(`if` 키워드와 같은 들여쓰기 수준에서) `print()` 호출과 f-string을 사용해 `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'`를 출력하세요.

그다음 줄에서 `eggs_to_use`를 `available_eggs`에서 빼서 남은 달걀 수를 `return`하세요.

# --hints--

`use_eggs` 함수 안에 `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'`를 출력하는 print 문이 있어야 합니다.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_stmt("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')")
`)
})
```

`eggs_to_use`를 `available_eggs`에서 빼서 남은 달걀 수를 반환해야 합니다.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_return("available_eggs - eggs_to_use")
`)
})
```

`print` 문 다음에 `available_eggs - eggs_to_use`를 반환해야 합니다.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs")
assert node.is_ordered("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')", "return available_eggs - eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

--fcc-editable-region--
    
--fcc-editable-region--
```
