---
id: 6a3e2b9b7eeca084dd4fb857
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

주방에 오신 것을 환영합니다! 요리사가 레시피에 의존해 음식을 준비하듯, 소프트웨어를 만들 때는 함수에 의존하게 됩니다. 이 워크숍에서는 재사용 가능한 Python 함수를 작성하는 방법과 범위를 마스터하여 코드에서 언제든지 어떤 재료(변수)가 사용 가능한지 정확히 이해하게 됩니다.

시작하려면 주방 재고를 위한 기본 데이터를 정의하세요. `available_eggs`, `available_flour`, `available_sugar`라는 세 변수를 만들고 각각에 `1`, `2`, `3` 값을 할당하세요.

# --hints--

`available_eggs`라는 변수를 정의하고 값으로 `1`을 할당하세요.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[0].is_equivalent("available_eggs = 1")
`)
})
```

`available_flour`라는 변수를 정의하고 값으로 `2`를 할당하세요.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_flour")[0].is_equivalent("available_flour = 2")
`)
})
```

`available_sugar`라는 변수를 정의하고 값으로 `3`을 할당하세요.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_sugar")[0].is_equivalent("available_sugar = 3")
`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
