---
id: 6a3e524fc3394ef52ed5b797
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

함수는 자체 지역 범위를 가집니다. 이는 함수 내부에서 생성된 변수가 그 주방의 저장실에 보관된 재료와 같아서, 특정 함수 내부에서만 접근할 수 있다는 뜻입니다.

하지만 함수는 완전히 고립되어 있지 않습니다. 함수는 자신의 범위 밖에 정의된 전역 변수를 읽을 수 있습니다. 그 변수들을 사용해 계산을 하거나, 결정을 내리거나, 새로운 지역 변수를 만들 수 있습니다.

`check_kitchen_stock` 함수 안에서 `pass` 키워드를 `total_items`라는 새 변수로 바꾸고, `available_eggs`, `available_flour`, `available_sugar`의 합을 할당하세요.

# --hints--

`check_kitchen_stock` 함수 안에 `total_items`라는 변수를 정의하세요.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_variable("total_items")
`)
})
```

`total_items` 변수는 `available_eggs`, `available_flour`, `available_sugar`의 합이어야 합니다.

```js
({
  test: () => runPython(`
variable = _Node(_code).find_function("check_kitchen_stock").find_variables("total_items")[0]
assert any(
  variable.is_equivalent(code)
  for code in [
    "total_items = available_eggs + available_flour + available_sugar",
    "total_items = available_flour + available_eggs + available_sugar",
    "total_items = available_sugar + available_flour + available_eggs",
    "total_items = available_eggs + available_sugar + available_flour",
    "total_items = available_flour + available_sugar + available_eggs",
    "total_items = available_sugar + available_eggs + available_flour"
  ]
)
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
--fcc-editable-region--
    pass
--fcc-editable-region--
```
