---
id: 6a3e6fd08125dfa615cb431b
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

`use_eggs` 함수가 완성되었으니 테스트해보세요. 매개변수가 필요한 함수를 호출할 때는 소괄호 안에 특정 값을 전달합니다. 이 값들을 인수라고 합니다.

다음은 인수를 사용해 함수를 호출하는 예입니다:

```py
def calculate_sum(a, b):
    print(a + b)
calculate_sum(3, 1) # Output: 4
```

코드 맨 아래에서 `use_eggs` 함수를 호출하고 두 개의 인수를 전달하세요: 전역 변수 `available_eggs`와 사용하려는 달걀 수를 나타내는 숫자 `1`입니다.

# --hints--

`use_eggs` 함수를 호출하고 `available_eggs`와 `1`을 인수로 전달하세요.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("use_eggs(available_eggs, 1)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

--fcc-editable-region--

--fcc-editable-region--
```
