---
id: 6a3e5ccb99fd6053af9032a1
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

주방에 무엇이 있는지 알게 되었으니, 그 재료들을 사용해 요리할 수 있습니다. 이를 위해 매개변수를 받는 새 함수를 만드세요. 매개변수는 함수를 호출할 때 전달하는 데이터를 위한 자리 표시자라고 생각하세요.

다음은 두 개의 매개변수 `a`와 `b`를 받아 그 합을 터미널에 출력하는 함수입니다.

```py
def calculate_sum(a, b):
    print(a + b)

calculate_sum(3, 1) # Output: 4
```

`check_kitchen_stock` 함수 아래에 `use_eggs`라는 이름의 빈 함수를 정의하세요. 이 함수는 순서대로 `available_eggs`와 `eggs_to_use` 두 개의 매개변수를 받습니다.

만약 매개변수 이름을 전역 변수와 정확히 같은 이름(예: `available_eggs`)으로 지정하면, 함수는 전역 변수 대신 자신의 지역 매개변수 `available_eggs`를 사용합니다. 이를 이름 가리기(name shadowing)라고 하며, 지역 변수가 전역 변수를 "가리는" 현상입니다.

함수 본문에는 일단 `pass` 키워드를 사용하세요.

# --hints--

`use_eggs`라는 이름의 함수를 정의해야 합니다.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("use_eggs")
`)
})
```

`use_eggs` 함수는 `available_eggs`와 `eggs_to_use`를 매개변수로 가져야 합니다.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_args("available_eggs, eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

--fcc-editable-region--

--fcc-editable-region--
```
