---
id: 6a3e70b94f67d9ee3ecd8486
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

방금 `use_eggs(available_eggs, 1)`를 호출해서 달걀 한 개를 사용했어요. 그러니까 이제 부엌에 `0`개의 달걀이 남아 있어야 하죠?

직접 확인해 보세요! 코드 맨 아래에서 `print()` 함수를 사용해 `available_eggs` 변수의 값을 출력하세요.

코드를 실행하면 놀라운 점을 발견할 거예요: 출력 결과가 여전히 `1`이에요! 왜 달걀 개수가 변하지 않았을까요? 다음 단계에서 이 미스터리를 탐구할 거예요.

# --hints--

코드 마지막에 `available_eggs` 변수를 출력하세요.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("print(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
--fcc-editable-region--

--fcc-editable-region--
```
