---
id: 6a3ea771b46a239c67628a07
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

`make_fried_egg` 함수는 이제 달걀이 있는지 확인하고, 사용하면 지역 달걀 개수를 업데이트합니다. 다음으로, 호출자가 사용할 수 있도록 그 개수를 반환하세요.

`make_fried_egg` 함수 맨 아래에서 `return` 키워드를 사용해 `available_eggs` 변수를 반환하세요. 이 줄이 `if`나 `else` 블록 안에 들여쓰기 되어 있지 않은지 확인하세요.

# --hints--

`make_fried_egg` 함수 끝에서 `available_eggs`를 반환해야 합니다.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").has_return("available_eggs")
`)
})
```

`if/else` 문 다음에 `available_eggs`를 반환해야 합니다.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg")
assert node.is_ordered(str(node.find_ifs()[0]), "return available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

--fcc-editable-region--
    
--fcc-editable-region--
```
