---
id: 6a3e71a0a3a44260e675b157
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

계란 개수가 변하지 않은 이유를 해결하고 실제로 아침 식사를 요리하려면 새 함수를 만드세요. 이 함수를 다른 함수를 호출해 작업을 수행하는 레시피로 생각하세요.

`print(available_eggs)` 문 아래에 `available_eggs`를 유일한 매개변수로 받는 빈 함수 `make_fried_egg`를 정의하세요.

함수 본문 안에는 일단 `pass` 키워드를 사용하세요.

# --hints--

`make_fried_egg`라는 이름의 함수를 정의해야 합니다.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("make_fried_egg")
`)
})
```

`make_fried_egg` 함수는 `available_eggs`를 매개변수로 가져야 합니다.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").has_args("available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

--fcc-editable-region--

--fcc-editable-region--
```
