---
id: 6a3e6ad6b9554e31dcc30cd6
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

달걀이 충분하지 않으면 요리하려고 하지 마세요! 함수가 더 이상 실행되지 않도록 멈추는 방법이 필요합니다.

이럴 때 `return` 키워드를 사용할 수 있습니다. 함수가 `return` 문을 만나면 즉시 실행이 정지되고 반환값을 호출된 곳으로 전달합니다.

예를 들어:

```py
def cook_dish():
    message = "The dish is ready to serve!"
    return message
    print("This line will never run.")

print(cook_dish())  # Output: The dish is ready to serve!
```

`if` 블록 안에서 출력문 아래에 `return` 키워드를 사용해 `available_eggs` 값을 반환하세요.

# --hints--

`if` 문 안에서 `available_eggs`를 반환해야 합니다.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").find_ifs()[0].has_return("available_eggs")
`)
})
```

`if` 블록 안에서 `print` 문 다음에 `available_eggs`를 반환해야 합니다.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs").find_ifs()[0]
assert node.is_ordered("print('The kitchen does not have enough eggs.')", "return available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
--fcc-editable-region--
        
--fcc-editable-region--
```
