---
id: 6a3ea8a55317ba7a08e21d5b
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

이제 이전에 했던 질문에 드디어 답할 수 있습니다: 왜 전역 계란 수가 `1`로 유지되었을까요? 함수가 값을 반환해도 기존 변수를 자동으로 덮어쓰지 않습니다. 새 값을 명시적으로 저장해야 합니다!

`make_fried_egg` 함수 안에서 `available_eggs` 변수를 사용하고 있지만, 코드 상단에 정의된 변수와는 다릅니다. 함수 안의 변수는 함수가 실행되는 동안에만 존재하는 지역 변수입니다. 함수가 끝나면 지역 변수는 사라지고 전역 `available_eggs`는 변하지 않습니다.

실제 주방 재고를 업데이트하려면, `make_fried_egg` 함수가 반환하는 새 계란 수로 전역 `available_eggs` 변수를 덮어쓰세요.

코드 맨 아래에서 `make_fried_egg` 함수를 호출할 때 `available_eggs`를 인수로 전달하고, 결과를 다시 전역 `available_eggs` 변수에 할당하세요.

# --hints--

`available_eggs` 변수를 `make_fried_egg(available_eggs)`의 결과로 다시 할당해야 합니다.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[-1].is_equivalent("available_eggs = make_fried_egg(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

--fcc-editable-region--

--fcc-editable-region--
```
