---
id: 6a3e58891e1ff13d07f4ab48
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

Python은 `def` 블록을 읽고 함수를 기억하지만, 그 안의 코드를 즉시 실행하지는 않습니다.

코드를 실행하고 결과를 보려면 함수를 **호출**(또는 실행)해야 합니다. 함수 이름 뒤에 소괄호를 붙여 입력하세요.

예를 들어:

```py
my_function()
```

코드 맨 아래, 함수 밖에서 `check_kitchen_stock` 함수를 호출하세요.

# --hints--

코드 끝에서 `check_kitchen_stock` 함수를 호출해야 합니다. 들여쓰기가 없도록 해서 전역에서 실행되게 하세요.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("check_kitchen_stock()")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

--fcc-editable-region--

--fcc-editable-region--
```
