---
id: 6a3e4ae26f2b7e4b5b026cf8
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

주방에 있는 것을 확인하는 함수를 만들 시간입니다. Python에서는 `def` 키워드 다음에 함수 이름, 소괄호 `()`, 그리고 콜론 `:`을 사용해 함수를 정의합니다.

간단한 함수를 정의하는 예시는 다음과 같습니다:

```py
def my_function():
    pass
```

`check_kitchen_stock`라는 이름의 빈 함수를 정의하세요. Python이 `SyntaxError`를 발생시키지 않도록 함수 본문 안에 `pass` 키워드를 사용하세요.

Python은 코드 블록을 정의할 때 들여쓰기에 의존한다는 점을 기억하세요. 함수 본문을 반드시 들여쓰기 하세요.

# --hints--

`check_kitchen_stock`라는 함수를 정의해야 합니다.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("check_kitchen_stock")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

--fcc-editable-region--

--fcc-editable-region--
```
