---
id: 6a3e7284e89ce3367e709331
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

계란 프라이를 만들기 전에 필요한 재료가 있는지 먼저 확인하세요. 레시피에 필요한 계란 수량과 현재 가지고 있는 계란 수량을 비교해서 확인할 수 있습니다.

`make_fried_egg` 함수 안에서 `pass` 키워드를 `has_enough_eggs`라는 변수로 바꾸세요. 그리고 `available_eggs`가 `1` 이상인지 확인한 결과를 할당하세요.

# --hints--

`has_enough_eggs`에 `available_eggs`가 `1` 이상인지 확인한 결과를 할당하세요.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg").find_variables("has_enough_eggs")[0]
assert node.is_equivalent("has_enough_eggs = available_eggs >= 1") or node.is_equivalent("has_enough_eggs = 1 <= available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
