---
id: 6a3e5475c35dbc138d77c6f4
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

총합을 계산했으니 재고를 출력하세요. 이전 수업에서 `print()` 함수와 **f-strings**를 사용해 변수 값을 출력에 포함할 수 있다는 것을 배웠습니다.

`check_kitchen_stock` 함수 안에서 `total_items` 변수 아래에 다음 네 개의 정확한 print 문을 추가하세요:

1. `print(f'The kitchen has {total_items} total items:')`
2. `print(f'- {available_eggs} eggs')`
3. `print(f'- {available_flour} flour')`
4. `print(f'- {available_sugar} sugar')`

# --hints--

`print(f'The kitchen has {total_items} total items:')` 문이 함수 안에 있어야 합니다.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'The kitchen has {total_items} total items:')")
`)
})
```

`print(f'The kitchen has {total_items} total items:')` 문이 함수 안에서 첫 번째 `print()` 호출이어야 합니다.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").find_calls("print")[0].is_equivalent("print(f'The kitchen has {total_items} total items:')")
`)
})
```

`print(f'- {available_eggs} eggs')` 문이 함수 안에 있어야 합니다.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_eggs} eggs')")
`)
})
```

`print(f'- {available_flour} flour')` 문이 함수 안에 있어야 합니다.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_flour} flour')")
`)
})
```

`print(f'- {available_sugar} sugar')` 문이 함수 안에 있어야 합니다.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_sugar} sugar')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
--fcc-editable-region--
    
--fcc-editable-region--
```
