---
id: 6a3e67fc5ef4045201c28219
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

`use_eggs` 함수 안에서 요리에 사용할 계란이 충분한지 확인해야 합니다. 이는 `if` 문을 사용해 할 수 있습니다.

`pass` 키워드를 `if` 문으로 바꾸어 `eggs_to_use`가 `available_eggs`보다 큰지(`>`) 확인하세요. `if` 블록 안에서 문자열 `The kitchen does not have enough eggs.`를 출력하세요.

# --hints--

`eggs_to_use > available_eggs`인지 확인하는 `if` 문이 있어야 합니다.

```js
({
  test: () => runPython(`
  cond = _Node(_code).find_function("use_eggs").find_ifs()[0].find_conditions()[0]
  assert cond.is_equivalent("eggs_to_use > available_eggs") or cond.is_equivalent("available_eggs < eggs_to_use")`)
})
```

`if` 문 안에서 `The kitchen does not have enough eggs.`를 출력해야 합니다.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").find_ifs()[0].has_call("print('The kitchen does not have enough eggs.')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
