---
id: 6a3eaa2caac22f67811fe40f
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

결정적인 순간입니다! `make_fried_egg` 함수가 반환한 새 값으로 전역 `available_eggs` 변수를 성공적으로 업데이트했습니다.

계란 수가 실제로 줄었는지 확인하려면 주방 재고를 마지막으로 점검하세요.

코드 맨 끝에서 `check_kitchen_stock()` 함수를 다시 호출하세요. 코드를 실행하면 주방에 이제 `0`개의 계란이 있음을 확인할 수 있을 것입니다!

주방 재고 추적 워크숍을 성공적으로 완료한 것을 축하합니다!

# --hints--

코드 끝에서 `check_kitchen_stock()` 함수를 두 번째로 호출해야 합니다.

```js
({
  test: () => runPython(`
node = _Node(_code)
assert len(node.find_calls("check_kitchen_stock")) == 2
assert node.is_ordered("available_eggs = make_fried_egg(available_eggs)", "check_kitchen_stock()")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

available_eggs = make_fried_egg(available_eggs)

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

available_eggs = make_fried_egg(available_eggs)

check_kitchen_stock()
```
