---
id: 6a3e73eabddd614f91cc235a
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

이제 앞서의 미스터리를 해결할 차례입니다! `use_eggs` 함수가 남은 달걀 수를 반환했을 때, 그 새 값을 어디에도 저장하지 않았습니다. 그냥 사라져 버렸어요!

재고를 실제로 업데이트하려면, 반환된 값을 변수에 할당해 저장해야 합니다.

`make_fried_egg` 함수 안에서 `if/else` 문을 사용해 `has_enough_eggs` 변수를 확인하세요.

- `has_enough_eggs`가 `True`라면, `available_eggs` 변수를 `use_eggs(available_eggs, 1)`의 결과로 다시 할당하고 `Made a fried egg. Yummy!` 메시지를 출력하세요.

- 그렇지 않으면 `else` 블록 안에서 `Could not make a fried egg. Not enough eggs!` 메시지를 출력하세요.

# --hints--

`if` 문이 있어야 하며 조건으로 `has_enough_eggs`를 사용하세요.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_conditions()[0].is_equivalent("has_enough_eggs")
`)
})
```

`if` 블록 안에서 `available_eggs`를 `use_eggs(available_eggs, 1)`의 결과로 다시 할당해야 합니다.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].find_variables("available_eggs")[0].is_equivalent("available_eggs = use_eggs(available_eggs, 1)")
`)
})
```

`has_enough_eggs`가 `True`일 때 `Made a fried egg. Yummy!` 메시지를 출력해야 합니다.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].has_call("print('Made a fried egg. Yummy!')")
`)
})
```

`has_enough_eggs`가 `False`일 때 `Could not make a fried egg. Not enough eggs!` 메시지를 출력하는 `else` 블록이 있어야 합니다.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg").find_if("has_enough_eggs")
conditions = node.find_conditions()
assert len(conditions) == 2 and conditions[1].is_empty()
assert node.find_bodies()[1].has_call("print('Could not make a fried egg. Not enough eggs!')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

--fcc-editable-region--
    
--fcc-editable-region--
```
