---
id: 68e4276483b0c71658aab1d3
title: الخطوة 14
challengeType: 1
dashedName: step-14
---

# --description--

كما تتذكر من الدروس وورش العمل السابقة، يمكنك استخراج أجزاء من سلسلة نصية باستخدام الطريقة `slice()`.

إليك تذكير:

```js
const str = "abcd";
str.slice(0, 3); // "abc"
```

يجب أن تغير التعيين الحالي لـ `camelCasedVersion` باستبدال السلسلة النصية الفارغة بـ `lowercaseWord.slice(0, 5)`.

الآن يجب أن ترى الكلمة `"camel"` في وحدة التحكم.

# --hints--

يجب ألا تكون قيمة المتغير `camelCasedVersion` سلسلة نصية فارغة.

```js
assert.isNotEmpty(camelCasedVersion);
```

يجب أن تعيّن `lowercaseWord.slice(0, 5)` إلى المتغير `camelCasedVersion`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5));
```

لا يجب أن تعيّن السلسلة النصية الحرفية `"camel"` إلى `camelCasedVersion`. يرجى إعادة قراءة التعليمات للحصول على الطريقة الصحيحة.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camel(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = ""
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
