---
id: 68e4276483b0c71658aab1ca
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستتدرب على العمل مع طرق مختلفة للسلاسل النصية تُستخدم لتنسيق السلاسل النصية.

ابدأ بإنشاء متغير اسمه `userInput` وقم بتعيين القيمة التالية له: `"   Hello World!   "`.

# --hints--

يجب أن يكون لديك متغير `userInput`.

```js
assert.exists(userInput);
```

يجب أن يكون متغير `userInput` سلسلة نصية.

```js
assert.isString(userInput);
```

يجب أن تعيّن `"   Hello World!   "` إلى متغير `userInput`. تأكد من وجود ثلاث مسافات قبل الحرف `H` وثلاث مسافات بعد العلامة `!`.

```js
assert.equal(userInput, "   Hello World!   ");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
