---
id: 68e4276483b0c71658aab1d2
title: الخطوة 13
challengeType: 1
dashedName: step-13
---

# --description--

ستحتاج أحيانًا إلى استخدام عدة طرق لإنشاء سلسلة نصية جديدة. في الخطوات التالية، ستستخدم بعض الطرق التي تعلمتها حتى الآن لإنشاء سلسلة نصية جديدة تستخدم اصطلاح تسمية تنسيق سنام الجمل.

ابدأ بإنشاء متغير يسمى `lowercaseWord` وقم بتعيين السلسلة النصية `"camelcase"` له.

أسفل متغير `lowercaseWord`، أنشئ متغيرًا آخر يسمى `camelCasedVersion` وقم بتعيين سلسلة نصية فارغة له مؤقتًا.

أسفل متغير `camelCasedVersion`، أضف دالة `console.log()` مع السلسلة النصية `"Camel cased version:"` تليها دالة `console.log()` أخرى للمتغير `camelCasedVersion`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن يكون لديك متغير يسمى `lowercaseWord`.

```js
assert.exists(lowercaseWord);
```

يجب أن يكون `lowercaseWord` سلسلة نصية.

```js
assert.isString(lowercaseWord);
```

يجب أن تعيّن السلسلة النصية `"camelcase"` إلى متغير `lowercaseWord`.

```js
assert.equal(lowercaseWord, "camelcase");
```

يجب أن يكون لديك متغير `camelCasedVersion`.

```js
assert.exists(camelCasedVersion);
```

يجب أن يكون متغير `camelCasedVersion` سلسلة نصية.

```js
assert.isString(camelCasedVersion);
```

يجب أن تعيّن سلسلة نصية فارغة إلى متغير `camelCasedVersion`.

```js
assert.isEmpty(camelCasedVersion);
```

يجب أن تسجل السلسلة النصية `"Camel cased version:"` في وحدة التحكم.

```js
assert.equal(getLogs()[12], "Camel cased version:");
```

يجب أن تسجل المتغير `camelCasedVersion` في وحدة التحكم.

```js
assert.equal(getLogs()[13], camelCasedVersion);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
