---
id: 68e4276483b0c71658aab1cc
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

قيمة `userInput` تحتوي على مسافات إضافية في بداية ونهاية السلسلة النصية. لإنشاء سلسلة نصية جديدة تُزال منها هذه الفراغات، يمكنك استخدام الطريقة `trim()`.

إليك مثالاً:

```js
const sentence = "  Learning about the trim method.  ";
sentence.trim(); // "Learning about the trim method."
```

أنشئ متغيرًا باسم `cleanedInput` وقم بتعيين نتيجة استخدام الطريقة `trim()` على المتغير `userInput` إليه.

# --hints--

يجب أن يكون لديك متغير `cleanedInput`.

```js
assert.exists(cleanedInput);
```

يجب أن يكون المتغير `cleanedInput` سلسلة نصية.

```js
assert.isString(cleanedInput);
```

يجب أن تُعيّن `userInput.trim()` إلى المتغير `cleanedInput`.

```js
assert.equal(cleanedInput, userInput.trim());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);
--fcc-editable-region--

--fcc-editable-region--
```
