---
id: 68e4276483b0c71658aab1d0
title: الخطوة 10
challengeType: 1
dashedName: step-10
---

# --description--

ستحتاج أحيانًا إلى إنشاء سلسلة نصية جديدة بحالة أحرف مختلفة.

لإنشاء سلسلة نصية جديدة مع تحويل جميع الحروف إلى أحرف كبيرة، يمكنك استخدام الطريقة `toUpperCase()`.

إليك مثالًا:

```js
const sentence = "I love to code!";
sentence.toUpperCase(); // "I LOVE TO CODE!"
```

أنشئ متغيرًا اسمه `upperCaseInput` وقم بتعيين نتيجة استخدام الطريقة `toUpperCase()` على المتغير `cleanedInput` له.

# --hints--

يجب أن يكون لديك متغير `upperCaseInput`.

```js
assert.exists(upperCaseInput);
```

يجب أن يكون المتغير `upperCaseInput` سلسلة نصية.

```js
assert.isString(upperCaseInput);
```

يجب تعيين `cleanedInput.toUpperCase()` إلى المتغير `upperCaseInput`.

```js
assert.equal(upperCaseInput, cleanedInput.toUpperCase());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

--fcc-editable-region--

--fcc-editable-region--
```
