---
id: 68e4276483b0c71658aab1ce
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

ستكون هناك أوقات تريد فيها فقط تقليم الفراغات في بداية السلسلة النصية. هنا تأتي فائدة طريقة `trimStart()`.

إليك مثالاً:

```js
const sentence = "  I love learning.  ";
sentence.trimStart(); // "I love learning.  "
```

تذكر أن السلاسل النصية غير قابلة للتغيير. هذه الطريقة تُنشئ سلسلة نصية جديدة بالنتيجة، ولا تُعدل الأصلية.

أنشئ متغيرًا باسم `trimmedStart` وقم بتعيين نتيجة استخدام طريقة `trimStart()` على المتغير `userInput` له.

# --hints--

يجب أن يكون لديك متغير `trimmedStart`.

```js
assert.exists(trimmedStart);
```

يجب أن يكون المتغير `trimmedStart` سلسلة نصية.

```js
assert.isString(trimmedStart);
```

يجب أن تعيّن `userInput.trimStart()` إلى متغير `trimmedStart`.

```js
assert.equal(trimmedStart, userInput.trimStart());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);
--fcc-editable-region--

--fcc-editable-region--
```
