---
id: 68e4276483b0c71658aab1d1
title: الخطوة 12
challengeType: 1
dashedName: step-12
---

# --description--

إذا كنت بحاجة إلى إنشاء سلسلة نصية جديدة تكون كل حروفها أحرف صغيرة، يمكنك استخدام الطريقة `toLowerCase()`.

ابدأ بإنشاء متغير يسمى `lowerCaseInput` وقم بتعيين نتيجة استخدام الطريقة `toLowerCase()` على المتغير `cleanedInput` له.

ثم سجّل السلسلة النصية `"Result of using the toLowerCase() method:"`.

أخيرًا، أضف استدعاء `console.log()` آخر للمتغير `lowerCaseInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن يكون لديك متغير `lowerCaseInput`.

```js
assert.exists(lowerCaseInput);
```

يجب أن يكون المتغير `lowerCaseInput` سلسلة نصية.

```js
assert.isString(lowerCaseInput);
```

يجب تعيين `cleanedInput.toLowerCase()` إلى المتغير `lowerCaseInput`.

```js
assert.equal(lowerCaseInput, cleanedInput.toLowerCase());
```

يجب أن تسجل السلسلة النصية `"Result of using the toLowerCase() method:"` في وحدة التحكم.

```js
assert.equal(getLogs()[10], "Result of using the toLowerCase() method:");
```

يجب أن تسجل المتغير `lowerCaseInput` في وحدة التحكم.

```js
assert.equal(getLogs()[11], lowerCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);


const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
