---
id: 68e4276483b0c71658aab1cd
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

حان الوقت الآن لرؤية النتيجة.

ابدأ بإضافة `console.log()` مع السلسلة النصية `"Result of trimming whitespace from both ends:"`.

ثم أسفل ذلك `console.log()`، أضف `console.log()` آخر مع المتغير `cleanedInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن تسجل السلسلة النصية `"Result of trimming whitespace from both ends:"` في وحدة التحكم.

```js
assert.equal(getLogs()[2], "Result of trimming whitespace from both ends:");
```

يجب أن تسجل المتغير `cleanedInput` في وحدة التحكم.

```js
assert.equal(getLogs()[3], cleanedInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
--fcc-editable-region--

--fcc-editable-region--
```
