---
id: 68ed6079fe0e3e98b94a1b75
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

بعد ذلك، أضف دالة `console.log()` مع السلسلة النصية `"Original input:"`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن تسجل السلسلة النصية `"Original input:"` في وحدة التحكم.

```js
assert.equal(getLogs()[0], "Original input:");
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
--fcc-editable-region--

--fcc-editable-region--
```
