---
id: 68f1179eaea3bb59ba434386
title: الخطوة 8
challengeType: 1
dashedName: step-8
---

# --description--

لقص الفراغات في نهاية سلسلة نصية، يمكنك استخدام الطريقة `trimEnd()`.

إليك مثالاً:

```js
const sentence = "  I love learning.  ";
sentence.trimEnd(); // "  I love learning."
```

أنشئ متغيرًا اسمه `trimmedEnd` وقم بتعيين نتيجة استخدام الطريقة `trimEnd()` على المتغير `userInput`.

# --hints--

يجب أن يكون لديك متغير `trimmedEnd`.

```js
assert.exists(trimmedEnd);
```

يجب أن يكون المتغير `trimmedEnd` سلسلة نصية.

```js
assert.isString(trimmedEnd);
```

يجب تعيين `userInput.trimEnd()` إلى المتغير `trimmedEnd`.

```js
assert.equal(trimmedEnd, userInput.trimEnd());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

--fcc-editable-region--

--fcc-editable-region--
```
