---
id: 68f3f5ca1d2a490c97f14f03
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

حان الوقت الآن لرؤية نتائج إزالة الفراغات في نهاية السلسلة النصية.

أسفل المتغير `trimmedEnd`، أضف دالة `console.log()` مع السلسلة النصية `"After using the trimEnd() method, trailing spaces removed:"`.

ثم أسفل ذلك `console.log()`، أضف دالة `console.log()` أخرى مع المتغير `trimmedEnd`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن تسجل السلسلة النصية `"After using the trimEnd() method, trailing spaces removed:"` في وحدة التحكم.

```js
assert.equal(getLogs()[6], "After using the trimEnd() method, trailing spaces removed:");
```

يجب أن تسجل المتغير `trimmedEnd` في وحدة التحكم.

```js
assert.equal(getLogs()[7], trimmedEnd);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
--fcc-editable-region--

--fcc-editable-region--
```
