---
id: 68f3eeaa65614107ecd5f1f2
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

أسفل أول `console.log()`، أضف `console.log()` آخر لتسجيل المتغير `userInput` في وحدة التحكم.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب تسجيل المتغير `userInput` في وحدة التحكم.

```js
assert.equal(getLogs()[1], userInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
--fcc-editable-region--

--fcc-editable-region--
```
