---
id: 68f3f3d65d51d80ac0b412fc
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

حان الوقت الآن لرؤية نتائج إزالة الفراغات في بداية السلسلة النصية.

أسفل متغير `trimmedStart`، أضف دالة `console.log()` مع السلسلة النصية `"After using the trimStart() method, leading spaces removed:"`.

ثم أسفل ذلك `console.log()`، أضف دالة `console.log()` أخرى مع المتغير `trimmedStart`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن تسجل السلسلة النصية `"After using the trimStart() method, leading spaces removed:"` في وحدة التحكم.

```js
assert.equal(getLogs()[4], "After using the trimStart() method, leading spaces removed:");
```

يجب أن تسجل المتغير `trimmedStart` في وحدة التحكم.

```js
assert.equal(getLogs()[5], trimmedStart);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
--fcc-editable-region--

--fcc-editable-region--
```
