---
id: 68f3f890010c870e17bbe754
title: الخطوة 11
challengeType: 1
dashedName: step-11
---

# --description--

حان الوقت الآن لرؤية النتائج.

أسفل متغير `upperCaseInput`، أضف دالة `console.log()` مع السلسلة النصية `"Result of using the toUpperCase() method:"`.

ثم أسفل تلك الدالة `console.log()`، أضف دالة `console.log()` أخرى مع المتغير `upperCaseInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن تسجل السلسلة النصية `"Result of using the toUpperCase() method:"` في وحدة التحكم.

```js
assert.equal(getLogs()[8], "Result of using the toUpperCase() method:");
```

يجب أن تسجل المتغير `upperCaseInput` في وحدة التحكم.

```js
assert.equal(getLogs()[9], upperCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
--fcc-editable-region--

--fcc-editable-region--
```
