---
id: 68f405c8c9df44132c8fe1bb
title: الخطوة 15
challengeType: 1
dashedName: step-15
---

# --description--

كما تتذكر من الدروس السابقة، يبدأ اصطلاح تسمية camelCase بحرف صغير للكلمة الأولى، وتبدأ كل كلمة لاحقة بحرف كبير يتبعه أحرف صغيرة.

إليك بعض الأمثلة التذكيرية:

```md
freeCodeCamp
camelCase
learningIsFun
```

الكلمة الثانية في المتغير `lowercaseWord` هي `"case"`. للوصول إلى الحرف `c` في تلك الكلمة، يمكنك استخدام `lowercaseWord[5]`.

حدّث تعيين `camelCasedVersion` لاستخدام عامل `+` لربط `lowercaseWord.slice(0, 5)` مع نتيجة استخدام الطريقة الصحيحة لتحويل السلاسل النصية إلى أحرف كبيرة على `lowercaseWord[5]`.

الآن يجب أن ترى `camelC` في وحدة التحكم.

# --hints--

يجب أن تربط `lowercaseWord.slice(0, 5)` مع `lowercaseWord[5].toUpperCase()`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase());
```

لا يجب أن تعيّن السلسلة النصية الحرفية `"camelC"` إلى `camelCasedVersion`. يرجى إعادة قراءة التعليمات للحصول على الطريقة الصحيحة.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelC(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5)
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

