---
id: 68f40925d810ae14b2d5a169
title: الخطوة 16
challengeType: 1
dashedName: step-16
---

# --description--

في الوقت الحالي، متغير `camelCasedVersion` يحمل القيمة `camelC`. لكن النتيجة المطلوبة هي `camelCase`.

للحصول على الحروف المتبقية `ase`، يمكنك استخدام الطريقة `slice()` مع قيمة فهرس سالبة. تذكر أن استخدام قيم الفهرس السالبة يخبر الطريقة `slice()` باستخراج النص من النهاية.

إليك مثال تذكيري:

```js
const str = "abcd";
str.slice(-2); // cd
```

حدّث تعيين `camelCasedVersion` ليشمل أيضًا دمج `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` مع نتيجة استخدام الطريقة `slice()` على `lowercaseWord`. مرر الرقم `-3` للطريقة `slice()`.

وبهذا التغيير الأخير، تكتمل ورشة العمل الخاصة بك!

# --hints--

يجب دمج `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` مع `lowercaseWord.slice(-3)`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase() + lowercaseWord.slice(-3));
```

لا يجب تعيين السلسلة النصية الحرفية `"camelCase"` إلى `camelCasedVersion`. يرجى إعادة قراءة التعليمات للحصول على الطريقة الصحيحة.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelCase(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

# --solutions--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";

const camelCasedVersion =
  lowercaseWord.slice(0, 5) +
  lowercaseWord[5].toUpperCase() +
  lowercaseWord.slice(-3);

console.log("Camel cased version:");
console.log(camelCasedVersion);
```
