---
id: 668f08ea07b99b1f4a91acab
title: بناء صفحة وصفة
challengeType: 25
dashedName: build-a-recipe-page
demoType: onClick
---

# --description--

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تحتوي على إعلان `!DOCTYPE html`.
1. يجب أن تحتوي على عنصر `html` مع تعيين `lang` إلى `en`.
1. يجب أن تحتوي على عنصر `head` يحتوي على عنصر `title` باسم وصفتك، وعنصر `meta` مع خاصية `charset` مضبوطة على `UTF-8`.
1. يجب أن تحتوي على عنصر `body`.
1. يجب أن تحتوي على عنصر `h1` باسم وصفتك.
1. يجب أن تحتوي على عنصر `p` يقدم الوصفة أسفل `h1`.
1. يجب أن تحتوي على عنصر `h2` واحد مع النص `Ingredients` لقسم المكونات.
1. يجب أن تحتوي على قائمة غير مرتبة (`ul` عنصر) بها أربعة عناصر قائمة (`li` عناصر) على الأقل تسرد مكوناتك أسفل أول عنصر `h2`.
1. يجب أن تحتوي على عنصر `h2` ثاني مع النص `Instructions` لقسم التعليمات.
1. يجب أن تحتوي على قائمة مرتبة (`ol` عنصر) بها أربعة عناصر قائمة على الأقل تسرد خطوات الوصفة بالترتيب، أسفل العنصر الثاني `h2`.
1. يجب أن تحتوي على عنصر `img` واحد مع خاصية `src` مضبوطة على صورة صالحة، يمكنك استخدام `https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg` إذا أردت، وخاصية `alt` تصف الصورة.

# --hints--

يجب أن تحتوي صفحة وصفتك على إعلان `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

يجب أن تحتوي على عنصر `html` مع تعيين `lang` إلى `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/gi);
```

يجب أن تحتوي على عنصر `head` داخل عنصر `html`.

```js
assert.match(code, /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i);
```

يجب أن تحتوي على عنصر `title` داخل عنصر `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i);
```

يجب أن يحتوي عنصر `title` على عنوان وصفتك.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

يجب أن تحتوي على عنصر `meta` داخل عنصر `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*meta[\s\S]*>[\s\S]*<\/\s*head\s*>/i);
```

يجب أن تحتوي خاصية `charset` في عنصر `meta` على القيمة `UTF-8`.

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

يجب أن تحتوي على عنصر `body` داخل عنصر `html`.

```js
assert.match(code, /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i);
```

يجب أن تحتوي على عنصر `h1` باسم وصفتك.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

يجب أن يكون لديك عنصر `h1` واحد فقط.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

يجب أن تحتوي على عنصر `p` أسفل عنصر `h1`.

```js
assert.strictEqual(document.querySelector('h1')?.nextElementSibling, document.querySelector('p'));
```

يجب أن يصف أول عنصر `p` وصفتك.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

يجب أن يحتوي أول عنصر `h2` على النص `Ingredients`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Ingredients');
```

يجب أن تحتوي على قائمة غير مرتبة أسفل أول عنصر `h2`.

```js
assert.strictEqual(document.querySelector('ul')?.previousElementSibling.tagName, 'H2');
```

يجب أن تحتوي القائمة غير المرتبة على أربعة عناصر قائمة على الأقل مع المكونات.

```js
const els = document.querySelectorAll('ul > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

يجب أن يحتوي العنصر الثاني `h2` على النص `Instructions`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Instructions');
```

يجب أن تحتوي على قائمة مرتبة أسفل العنصر الثاني `h2`.

```js
assert.strictEqual(document.querySelectorAll('h2')?.[1]?.nextElementSibling?.tagName, "OL");
```

يجب أن تحتوي القائمة المرتبة على أربعة عناصر قائمة على الأقل مع التعليمات.

```js
const els = document.querySelectorAll('ol > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

يجب أن تحتوي على عنصر `img` واحد على الأقل.

```js
assert.exists(document.querySelector('img'));
```

يجب أن تحتوي جميع عناصر `img` على خاصية `src` صالحة وقيمة صحيحة.

```js
const img = document.querySelector('img');
const rawSrc = img?.getAttribute('src');
const resolvedSrc = img?.src;
const re = new RegExp(window.location.href, "ig");

assert.isAbove(rawSrc?.trim().length, 0, "The 'src' attribute must be explicitly set.");
assert.notMatch(resolvedSrc, re, "The 'src' should not start with the current page URL");

img.onload = () => {
  console.log('Image loaded successfully.');
};

img.onerror = (error) => {
  console.error('Image failed to load:', error);
  assert.fail("Your image's URL should be valid."); // Make the test instafail
};

if (img.complete) {
  img.onload && img.onload();
};
```

يجب أن تحتوي جميع عناصر `img` على خاصية `alt` لوصف الصورة.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Chocolate chip cookies recipe</title>
</head>

<body>
  <h1>Chocolate Chip Cookies</h1>
  <p>Welcome to the ultimate guide for making mini chocolate chip cookies! These bite-sized treats are perfect for
    satisfying your sweet tooth without overindulging. Follow this simple recipe to create delicious,
    crispy-on-the-outside, chewy-on-the-inside mini chocolate chip cookies that everyone will love.</p>
  <img src="https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg" alt="Ingredients for baking: three eggs, a bowl of flour, a glass of milk, and a whisk arranged on a wooden table.">
  <h2>Ingredients</h2>
  <ul>
    <li>1 cup all-purpose flour</li>
    <li>1/2 teaspoon baking soda</li>
    <li>1/4 cup unsalted butter, softened</li>
    <li>1/4 cup granulated sugar</li>
    <li>1/2 teaspoon vanilla extract</li>
    <li>1/2 cup mini chocolate chips</li>
  </ul>
  <h2>Instructions</h2>
  <ol>
    <li>Preheat your oven to 350°F (175°C) and line a baking sheet with parchment paper.</li>
    <li>In a bowl, whisk together the flour and baking soda.</li>
    <li>In another bowl, beat the butter, sugar, and vanilla extract until creamy.</li>
    <li>Gradually add the dry ingredients to the wet mixture, then fold in the mini chocolate chips.</li>
    <li>Drop small spoonfuls of dough onto the baking sheet.</li>
    <li>Bake for 8-10 minutes, then let cool before enjoying!</li>
  </ol>
</body>

</html>
```
