---
id: 684aaf9ec670c68d20efd0d0
title: 构建用户配置管理器
challengeType: 27
dashedName: build-a-user-configuration-manager
saveSubmissionToDB: true
---

# --description--

在本实验中，你将构建一个用户配置管理器，允许用户管理他们的设置，例如主题、语言和通知。你将实现用于添加、更新、删除和查看用户设置的函数。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该定义一个名为 `add_setting` 的函数，带有两个参数，分别表现为一个设置字典和一个包含密钥-值点对的元组。

1. `add_setting` 函数应当：
    
   - 将密钥和值转换为小写。
   - 如果密钥设置存在，返回 `Setting '[key]' already exists! Cannot add a new setting with this name.`
   - 如果密钥设置不存在，为给定的设置字典添加该键值点对并返回 `Setting '[key]' added with value '[value]' successfully!`。
   - 返回的消息应将密钥和值都转换为小写。

1. 你应该定义一个名为 `update_setting` 的函数，带有两个参数，分别表现为一个设置字典和一个包含密钥-值点对的元组。

1. `update_setting` 函数应：
       
   - 将密钥和值转换为小写。
   - 如果密钥设置存在，则更新给定设置字典中的其值并返回：`Setting '[key]' updated to '[value]' successfully!`
   - 如果密钥设置不存在，返回 `Setting '[key]' does not exist! Cannot update a non-existing setting.`
   - 返回的消息应将密钥和值都转换为小写。


1. 你应该定义一个名为 `delete_setting` 的函数，带有两个参数，表现为一个设置字典和一个密钥。

1. `delete_setting` 函数应当：
    
   - 将传入的密钥转换为小写。
   - 如果密钥设置存在，从给定的设置字典中移除该密钥-值点对并返回 `Setting '[key]' deleted successfully!`
   - 如果密钥设置不存在，返回 `Setting not found!`
   - 返回的消息应包含小写的密钥。

1. 你应该定义一个名为 `view_settings` 的函数，带有一个表现设置字典的参数。

1. `view_settings` 函数应：
   
   - 如果给定的设置字典为空，则返回 `No settings available.`。
   - 如果字典包含任何设置，返回一个显示设置的字串。该字串应以 `Current User Settings:` 开头，后跟密钥-值点对，每个点对占一行且密钥首字母大写。例如，`view_settings({'theme': 'dark', 'notifications': 'enabled', 'volume': 'high'})` 应返回：

   ```md
   Current User Settings:
   Theme: dark
   Notifications: enabled
   Volume: high
   
   ```

1. 为了测试代码，你应该创建一个名为 `test_settings` 的字典来保存一些用户配置偏好。

# --hints--

你应该创建一个名为 `test_settings` 的字典并为它添加一些值。

```js
({
    test: () => runPython(`
        is_dict = isinstance(test_settings, dict)
        length = len(test_settings)
        assert is_dict and length > 0
    `)
})
```

你应该定义一个名为 `add_setting` 的函数。

```js
({ test: () => assert(runPython(`_Node(_code).has_function("add_setting")`)) })
```

`add_setting` 函数应有两个参数。

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(add_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`add_setting` 应将密钥转换为小写。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("VOLUME", "high"))
        assert "volume" in settings_in_test
        assert "VOLUME" not in settings_in_test
    `)
})
```

`add_setting` 应将值转换为小写。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "HIGH"))
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting({'theme': 'light'}, ('THEME', 'dark'))` 应该返回错误信息 `Setting 'theme' already exists! Cannot add a new setting with this name.`。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = add_setting(settings_in_test, ("THEME", "dark"))
        expected_msg = "Setting 'theme' already exists! Cannot add a new setting with this name."
        assert result == expected_msg
    `)
})
```

`add_setting({'theme': 'light'}, ('volume', 'high'))` 应该添加一个新的密钥-值点对并返回成功消息 `Setting 'volume' added with value 'high' successfully!`。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' added with value 'high' successfully!"
        assert result == expected_msg
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting` 应正确地将给定的密钥-值点对添加到字典中。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("os", "Ubuntu"))
        assert "os" in settings_in_test
        assert settings_in_test["os"] == "ubuntu"
    `)
})
```

你应该定义一个名为 `update_setting` 的函数。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('update_setting')
    `)
})
```

`update_setting` 函数应有两个参数。

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(update_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`update_setting` 函数应将 key 转换为小写。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("THEME", "dark"))
        assert "THEME" not in settings_in_test
        assert "theme" in settings_in_test
        assert settings_in_test["theme"] == "dark"
    `)
})
```

`update_setting` 函数应将 value 转换为小写。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        # Use the existing "theme" key but an uppercase value
        result = update_setting(settings_in_test, ("theme", "HIGH"))
        assert settings_in_test["theme"] == "high"
    `)
})
```

`update_setting({'theme': 'light'}, ('theme', 'dark'))` 应该更新一个现有的密钥并返回成功消息 `Setting 'theme' updated to 'dark' successfully!`。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("theme", "dark"))
        expected_msg = "Setting 'theme' updated to 'dark' successfully!"
        assert result == expected_msg
        assert settings_in_test["theme"] == "dark"
    `)
})
```

`update_setting({'theme': 'light'}, ('volume', 'high'))` 应该在密钥不存在时返回错误信息 `Setting 'volume' does not exist! Cannot update a non-existing setting.`。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = update_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' does not exist! Cannot update a non-existing setting."
        assert result == expected_msg
    `)
})
```

`update_setting` 应正确更新字典中给定的密钥-值点对。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("language", "spanish"))
        assert "language" in settings_in_test
        assert settings_in_test["language"] == "spanish"
    `)
})
```

你应该定义一个名为 `delete_setting` 的函数。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('delete_setting')
    `)
})
```

`delete_setting` 函数应有两个参数。

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(delete_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`delete_setting` 应将密钥转换为小写。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "THEME")
        assert "theme" not in settings_in_test
        assert "THEME" not in settings_in_test
    `)
})
```

`delete_setting({'theme': 'light'}, 'theme')` 应删除现有的密钥并返回成功消息 `Setting 'theme' deleted successfully!`。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "theme")
        expected_msg = "Setting 'theme' deleted successfully!"
        assert result == expected_msg
        assert "theme" not in settings_in_test
    `)
})
```

`delete_setting` 应该在密钥不存在时返回错误信息 `Setting not found!`。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = delete_setting(settings_in_test, "volume")
        expected_msg = "Setting not found!"
        assert result == expected_msg
    `)
})
```

`delete_setting` 应正确地从字典中移除给定的密钥。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "language")
        assert "language" not in settings_in_test
    `)
})
```

你应该定义一个名为 `view_settings` 的函数。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('view_settings')
    `)
})
```

`view_settings` 函数应有一个参数。

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(view_settings)
        assert len(sig.parameters) == 1
    `)
})
```

`view_settings` 应该在给定字典为空时返回消息 `No settings available.`。

```js
({
    test: () => runPython(`
        empty_settings = {}
        result = view_settings(empty_settings)
        expected_msg = "No settings available."
        assert result == expected_msg
    `)
})
```

`view_settings` 应返回非空字典的格式化设置。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "notifications": "enabled", "volume": "high"}
        result = view_settings(settings_in_test)
        assert result.startswith("Current User Settings:")
        assert "Theme: dark" in result
        assert "Notifications: enabled" in result
        assert "Volume: high" in result
    `)
})
```

`view_settings` 应该将每个设置名称的首字母大写。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "volume": "high"}
        result = view_settings(settings_in_test)
        assert "Theme:" in result
        assert "Volume:" in result
        assert "theme:" not in result
        assert "volume:" not in result
    `)
})
```

`view_settings` 应显示正确的结果并以换行字符结尾。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English", "notifications": "enabled"}
        result = view_settings(settings_in_test)
        assert result == "Current User Settings:\\nTheme: light\\nLanguage: English\\nNotifications: enabled\\n"
    `)
})

```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def add_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        return f"Setting '{key}' already exists! Cannot add a new setting with this name."
    else:
        settings[key] = value
        return f"Setting '{key}' added with value '{value}' successfully!"

def update_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        settings[key] = value
        return f"Setting '{key}' updated to '{value}' successfully!"
    else:
        return f"Setting '{key}' does not exist! Cannot update a non-existing setting."

def delete_setting(settings, key):
    key = key.lower()
    if key in settings:
        del settings[key]
        return f"Setting '{key}' deleted successfully!"
    else:
        return "Setting not found!"

def view_settings(settings):
    if not settings:
        return "No settings available."
    else:
        settings_str = "Current User Settings:\n"
        for key, value in settings.items():
            settings_str += f"{key.capitalize()}: {value}\n"
        return settings_str

test_settings = {
    "theme": "light",
    "language": "English",
    "notifications": "enabled"
}

print(add_setting(test_settings, ("VOLUME", "HIGH")))
print(update_setting(test_settings, ("theme", "dark")))
print(delete_setting(test_settings, "language"))
print(view_settings(test_settings))
```
