---
id: 68de40f53fcaf106f51cc20f
title: Створіть HTML-плеєр для аудіо та відео
challengeType: 25
dashedName: build-an-html-audio-and-video-player
demoType: onClick
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні мати `h1` елемент для головного заголовка сторінки.
2. Ви повинні мати два `section` елементи.
3. Всередині першого `section` елемента має бути `h2` елемент для заголовка відтворюваного відео.
4. Під `h2` елементом має бути `video` елемент з атрибутами `controls` та `width`. Атрибут `width` має бути встановлений у `640`.
5. Всередині `video` елемента має бути `source` елемент з атрибутом `src`, що вказує на відеофайл, та атрибутом `type`.
   - Ви можете використати `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`.
6. Всередині другого `section` елемента має бути `h2` елемент для заголовка відтворюваної пісні.
7. Під `h2` елементом має бути `audio` елемент з атрибутами `controls` та `loop`, а також атрибутом `src`, що вказує на аудіофайл.
   - Ви можете використати `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`.
   - Або `https://cdn.freecodecamp.org/curriculum/js-music-player/we-are-going-to-make-it.mp3`.

# --hints--

Ви повинні мати `h1` елемент для головного заголовка сторінки.

```js
const h1 = document.querySelector('h1');
assert.exists(h1);
```

Ваш `h1` не повинен бути порожнім.

```js
const h1 = document.querySelector('h1');
assert.isNotEmpty(h1.textContent.trim());
```

Ви повинні мати лише один `h1` елемент.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Ви повинні мати два `section` елементи.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);

```

Всередині першого `section` елемента має бути `h2` елемент для заголовка відтворюваного відео.

```js
const section = document.querySelector('section');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

Під `h2` елементом має бути `video` елемент.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
assert.exists(video);
```

Елемент `video` має мати атрибут `controls`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const controlsAttr = video.controls;
assert.isTrue(controlsAttr);
```

Елемент `video` має мати атрибут `width`, встановлений у `640`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const widthAttr = video.getAttribute('width');
assert.equal(widthAttr, '640');
```

Всередині `video` елемента має бути `source` елемент.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
assert.exists(source);
```

Елемент `source` має мати атрибут `src`, що вказує на відеофайл.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const srcAttr = source.getAttribute('src');
assert.match(srcAttr, /\.(mp4|webm|ogg|avi|mov)$/i);
```

Елемент `source` має мати атрибут `type`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
assert.isNotEmpty(typeAttr);
```

Атрибут `type` має відповідати типу медіаконтенту розширення файлу атрибута `src`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
const srcAttr = source.getAttribute('src');
const srcFileExt = srcAttr.match(/[^\.]+$/g)[0];
const fileExtToMediaTypes = {
  mp4: ['video/mp4'],
  webm: ['video/webm'],
  ogg: ['video/ogg'],
  avi: ['video/avi', 'video/vnd.avi', 'video/msvideo', 'video/x-msvideo'],
  mov: ['video/quicktime']
};
const srcMediaTypes = fileExtToMediaTypes[srcFileExt];
assert.oneOf(typeAttr, srcMediaTypes);
```

Всередині другого `section` елемента має бути `h2` елемент для заголовка відтворюваної пісні.

```js
const section = document.querySelector('section:nth-of-type(2)');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

Усі ваші `h2` елементи мають містити текст.

```js
const h2s = document.querySelectorAll('h2');
assert.isNotEmpty(h2s);
h2s.forEach((h2) => assert.isNotEmpty(h2.textContent.trim()));
```

Під `h2` елементом має бути `audio` елемент.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
assert.exists(audio);
```

Елемент `audio` має мати атрибут `controls`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const controlsAttr = audio.controls;
assert.isTrue(controlsAttr);
```

Елемент `audio` має мати атрибут `loop`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const loopAttr = audio.loop;
assert.isTrue(loopAttr);
```

Елемент `audio` має мати атрибут `src`, що вказує на аудіофайл.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const srcAttr = audio.getAttribute('src');
assert.match(srcAttr, /\.(mp3|wav|ogg|aac|flac)$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>
  <main>
    <h1>HTML Audio and Video Lab</h1>

    <section>
      <h2>What is the map method and how does it work?</h2>
      <video controls width="640">
        <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4" />
      </video>
    </section>

    <section>
      <h2>Sailing Away</h2>
      <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" controls loop>
      </audio>
    </section>

  </main>
</body>
</html>
```
