---
id: 66a207974c806a19d6607073
title: 建置一本書庫存應用程式
challengeType: 25
dashedName: build-a-book-inventory-app
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `h1` 元素，裡面包含文字 `Book Inventory`。
1. 你應該有一個 `table` 元素，其行名稱為 `Title`、`Author`、`Category`、`Status` 和 `Rate`。
1. 你的表格應該至少有四列，第一列為行標題，其餘列填入書籍資訊。
1. 表格主體內的每一個表格列應該具有 `read`、`to-read` 或 `in-progress` 類別。
1. `Status` 行的 `td` 元素應包含一個 `class` 為 `status` 的 `span` 元素，包圍根據該列的類別而定的文字 `Read`、`To Read` 或 `In Progress`。
1. `Rate` 行的 `td` 元素應包含一個具有 `rate` `class` 的 `span` 元素，該元素包裹三個空的 `span` 元素。
1. 放置在 `.read` 列內的 `.rate` 元素應該有一個額外的類別，其值為 `one`、`two` 或 `three`，取決於個人評分。此值應該放在 `rate` 之後。
1. 你應該創建三個屬性選擇器來標的具有 `read`、`to-read` 和 `in-progress` 類別的元素，並將它們的 `background-image` 屬性設定為使用你選擇的 `linear-gradient`。
1. 你應該將每個 `span` 元素的 `display` 屬性設為 `inline-block`。
1. 你應該使用屬性選擇器來標的具有 `status` 類別且為具有 `to-read` 類別的 `tr` 元素的子孫的 `span` 元素，並設定它們的 `border` 和 `background-image` 屬性。
1. 你應該使用屬性選擇器來標的具有 `read` 類別的 `tr` 元素的子孫中，類別為 `status` 的 `span` 元素，並設定它們的 `border` 和 `background-image` 屬性。
1. 你應該使用屬性選擇器來標的具有 `status` 類別且為具有 `in-progress` 類別的 `tr` 元素的子孫的 `span` 元素，並設定它們的 `border` 和 `background-image` 屬性。
1. 你應該使用屬性選擇器來標的具有 `status` 類別的 `span` 元素以及類別值以 `rate` 開頭的 `span` 元素，並設定它們的 `height`、`width` 和 `padding` 屬性。
1. 你應該使用屬性選擇器來標的 `span` 元素，這些元素是 `class` 值以 `rate` 開頭的 `span` 元素的直接子元素，並設定它們的 `border`、`border-radius`、`margin`、`height`、`width` 和 `background-color` 屬性。
1. 你應該使用屬性選擇器來標的 `span` 元素中 `class` 值包含字 `one` 的第一個後代，並將其 `background-image` 屬性設定為使用 `linear-gradient`。
1. 你應該使用屬性選擇器來標的 `span` 元素中 `class` 值包含字 `two` 的前 2 個後代，並將它們的 `background-image` 屬性設定為使用 `linear-gradient`。
1. 你應該使用屬性選擇器來標的三個作為 `class` 值一部分且包含字 `three` 的 `span` 元素的子孫的 `span` 元素，並將它們的 `background-image` 屬性設定為使用 `linear-gradient`。

# --hints--

你應該有一個 `h1` 元素，裡面包含文字 `Book Inventory`。

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'Book Inventory');
```

你應該只有一個 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你應該有一個 `table` 元素。

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

你應該在 `table` 裡面有一個 `thead` 元素和一個 `tbody` 元素。

```js
assert.lengthOf(document.querySelectorAll("table > thead"), 1);
assert.lengthOf(document.querySelectorAll("table > tbody"), 1);
```

在 `thead` 內應該有一個包含 5 個 `th` 元素的 `tr`。

```js
assert.lengthOf(document.querySelectorAll("table > thead > tr"), 1);
assert.lengthOf(document.querySelectorAll("table > thead > tr > th"), 5);
```

你的第一行應該以文字 `Title` 作為標題。

```js
assert.equal(document.querySelectorAll('th')[0]?.textContent.trim(), 'Title');
```

你的第二個行應該以 `Author` 作為標題文字。

```js
assert.equal(document.querySelectorAll('th')[1]?.textContent.trim(), 'Author');
```

你的第三行應該以 `Category` 作為標題文字。

```js
assert.equal(document.querySelectorAll('th')[2]?.textContent.trim(), 'Category');
```

你的第四行標題應該是文字 `Status`。

```js
assert.equal(document.querySelectorAll('th')[3]?.textContent.trim(), 'Status');
```

你的第五行標題應該是文字 `Rate`。

```js
assert.equal(document.querySelectorAll('th')[4]?.textContent.trim(), 'Rate');
```

你的表格應該至少有四列。

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
```

每列應該始終有 5 行。

```js
const rows = document.querySelectorAll('tbody tr');
assert.isNotEmpty(rows);
for (const row of rows) {
  assert.lengthOf(row.children, 5);
}
```

除了標題列之外，每個資料表列都應該有 `read`、`to-read` 或 `in-progress` 類別。

```js
const rows = document.querySelectorAll('tr');
assert.isNotEmpty(rows);
for (let i = 1; i < rows.length; i++) {
  const classList = [...rows[i].classList];
  const currentProgress = classList[0]; 
  assert.oneOf(currentProgress, ['read', 'to-read', "in-progress"]); 
}
```

`Status` 行的 `td` 元素應包含一個 `span` 元素。

```js
const statusColumnData = document.querySelectorAll('td:nth-child(4)');
assert.isNotEmpty(statusColumnData);
for (let e of statusColumnData) {
  assert.equal(e?.children[0]?.tagName, 'SPAN');
}
```

`Status` 行的每個 `span` 元素都應該有 `status` 類別。

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
assert.isNotEmpty(statusSpans);
for (let e of statusSpans) {
    assert.isTrue(e.classList.contains('status'));
}
```

每個 `.status` 元素應該根據其列的類別，顯示 `Read`、`To Read` 或 `In Progress` 的文字。

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
const rows = Array.from(document.querySelectorAll('tr')).slice(1);
assert.isNotEmpty(statusSpans);
for (let i = 0; i < rows.length; i++) {
  switch (statusSpans[i]?.innerText.trim()) {
        case 'Read':
            assert.isTrue(rows[i].classList.contains('read'));
            break;
        case 'To Read':
            assert.isTrue(rows[i].classList.contains('to-read'));
            break;
        case 'In Progress':
            assert.isTrue(rows[i].classList.contains('in-progress'));
            break;
        default:
            assert.fail(); 
    }
}
```

`Rate` 行的 `td` 元素應包含一個 `span` 元素。

```js
const rateColumnData = document.querySelectorAll('tr td:last-child');
assert.isNotEmpty(rateColumnData);
for (let e of rateColumnData) {
    assert.equal(e.children[0]?.tagName, 'SPAN')
}
```

每個作為 `Rate` 行的 `td` 元素直接子元素的 `span` 元素，其第一個類別應為 `rate` 類別。

```js
const rateSpans = document.querySelectorAll('tr td:last-child > span:first-child');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.classList[0], 'rate');
}
```

每個 `.rate` 元素應包含三個空的 `span` 元素。

```js
const rateSpans = document.getElementsByClassName('rate');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.children.length, 3);
    for (let child of e.children) {
        assert.equal(child.tagName, 'SPAN');
        assert.equal(child.innerText.length, 0);
    }
}
```

放置在 `.read` 列中的 `.rate` 元素應在 `rate` 類別後面額外加上類別，該類別的值應為 `one`、`two` 或 `three`，取決於個人評分。

```js
const readBooksRates = document.querySelectorAll('.read .rate');
assert.isNotEmpty(readBooksRates);
for (let e of readBooksRates) {
    assert.oneOf(e.classList[1], ['one', 'two', 'three']);
}
```

你應該有一個屬性選擇器來標的具有 `read` 類別的列。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"]'));
```

你應該使用屬性選擇器來標的具有 `read` 類別的列，並將它們的 `background-image` 屬性設為你選擇的線性漸層。

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="read"]')?.backgroundImage, 'linear-gradient(');
```

你應該有一個屬性選擇器來標的具有 `to-read` 類別的列。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]'));
```

你應該使用屬性選擇器來標的具有 `to-read` 類別的列，並將它們的 `background-image` 屬性設為你選擇的線性漸層。

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]')?.backgroundImage, 'linear-gradient(');
```

你應該有一個屬性選擇器來標的具有 `in-progress` 類別的列。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]'));
```

你應該使用屬性選擇器來標的具有 `in-progress` 類別的列，並將它們的 `background-image` 屬性設為你選擇的線性漸層。

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]')?.backgroundImage, 'linear-gradient(');
```

你應該將每個 `span` 元素的 `display` 屬性設為 `inline-block`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('span')?.getPropVal('display'), 'inline-block');
```

你應該有一個屬性選擇器來標的具有 `to-read` 類別的 `tr` 元素的子孫中，類別為 `status` 的 `span` 元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]'));
```

你應該使用屬性選擇器來標的具有 `to-read` 類別的 `tr` 元素的子孫中，類別為 `status` 的 `span` 元素，並設定它們的 `border` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.border);
```

你應該使用屬性選擇器來標的具有 `to-read` 類別的 `tr` 元素的子孫中，具有 `status` 類別的 `span` 元素，並設定它們的 `background-image` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.backgroundImage);
```

你應該有一個屬性選擇器，用來標的具有 `read` 類別的 `tr` 元素的子孫中，具有 `status` 類別的 `span` 元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]'));
```

你應該使用屬性選擇器來標的具有 `read` 類別的 `tr` 元素的子孫中，類別為 `status` 的 `span` 元素，並設定它們的 `border` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.border);
```

你應該使用屬性選擇器來標的具有 `read` 類別的 `tr` 元素的子孫中，類別為 `status` 的 `span` 元素，並設定它們的 `background-image` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.backgroundImage);
```

你應該有一個屬性選擇器，以標的為具有 `in-progress` 類別的 `tr` 元素的後代且具有 `status` 類別的 `span` 元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]'));
```

你應該使用屬性選擇器來標的具有 `in-progress` 類別的 `tr` 元素的子孫中，類別為 `status` 的 `span` 元素，並設定它們的 `border` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.border);
```

你應該使用屬性選擇器來標的具有 `in-progress` 類別的 `tr` 元素的子孫中，類別為 `status` 的 `span` 元素，並設定它們的 `background-image` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.backgroundImage);
```

你應該有一個屬性選擇器來標的具有 `status` 類別的 `span` 元素以及類別值以 `rate` 開頭的 `span` 元素。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.exists(selector1 || selector2);
```

你應該使用屬性選擇器來標的具有 `status` 類別的 `span` 元素以及類別值以 `rate` 開頭的 `span` 元素，並設定它們的 `height` 屬性。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.height || selector2?.height);
```

你應該使用屬性選擇器來標的具有 `status` 類別的 `span` 元素以及類別值以 `rate` 開頭的 `span` 元素，並設定它們的 `width` 屬性。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.width || selector2?.width);
```

你應該使用屬性選擇器來標的具有 `status` 類別的 `span` 元素以及類別值以 `rate` 開頭的 `span` 元素，並設定它們的 `padding` 屬性。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.padding || selector2?.padding);
```

你應該有一個屬性選擇器來標的 `span` 元素，該元素是 `class` 值以 `rate` 開頭的 `span` 元素的直接子元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span'));
```

你應該使用屬性選擇器來標的 `span` 元素，這些元素是 `class` 值以 `rate` 開頭的 `span` 元素的直接子元素，並設定它們的 `border` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border'));
```

你應該使用屬性選擇器來標的 `span` 元素，這些元素是 `class` 值以 `rate` 開頭的 `span` 元素的直接子元素，並設定它們的 `border-radius` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border-radius'));
```

你應該使用屬性選擇器來標的 `span` 元素，這些元素是 `class` 值以 `rate` 開頭的 `span` 元素的直接子元素，並設定它們的 `margin` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('margin'));
```

你應該使用屬性選擇器來標的 `span` 元素，這些元素是 `class` 值以 `rate` 開頭的 `span` 元素的直接子元素，並設定它們的 `height` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('height'));
```

你應該使用屬性選擇器來標的 `span` 元素，這些元素是 `class` 值以 `rate` 開頭的 `span` 元素的直接子元素，並設定它們的 `width` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('width'));
```

你應該使用屬性選擇器來標的 `span` 元素，這些元素是 `class` 值以 `rate` 開頭的 `span` 元素的直接子元素，並設定它們的 `background-color` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('background-color'));
```

你應該有一個屬性選擇器，以標的為 `span` 元素中 `class` 值包含字 `one` 的第一個後代。

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

你應該使用屬性選擇器來標的 `span` 元素中 `class` 值包含字 `one` 的第一個後代，並將其 `background-image` 屬性設定為使用 `linear-gradient`。

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image').includes('linear-gradient('));
```

你應該有一個屬性選擇器，用來標的 `span` 元素中 `class` 值包含字 `two` 的前 2 個後代。

```js

const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

你應該使用屬性選擇器來標的 `span` 元素中 `class` 值包含字 `two` 的前 2 個後代，並將它們的 `background-image` 屬性設定為使用 `linear-gradient`。

```js
const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]

assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.backgroundImage.includes('linear-gradient(')) 
```

你應該有一個屬性選擇器來標的那些 `span` 元素，這些元素是 `span` 元素的後代，而這些 `span` 元素的 `class` 值中包含字 `three`。

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.exists(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

你應該使用屬性選擇器來標的那些 `span` 元素，這些元素是 `span` 元素的後代，而該 `span` 元素的 `class` 值中包含字 `three`，並將它們的 `background-image` 屬性設定為使用 `linear-gradient`。

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.include(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image'), 'linear-gradient(');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main>
    <h1>Book Inventory</h1>
    <div class="table-container">
      <table>
        <thead>
          <tr>
            <th>Title</th>
            <th>Author</th>
            <th>Category</th>
            <th>Status</th>
            <th>Rate</th>
          </tr>
        </thead>
        <tbody>
          <tr class="read">
            <td>The Three Musketeers</td>
            <td>Alexandre Dumas</td>
            <td>Historical Novel</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate three">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="to-read">
            <td>The Plague</td>
            <td>Albert Camus</td>
            <td>Philosophical Novel</td>
            <td><span class="status">To Read</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>The Metamorphosis</td>
            <td>Franz Kafka</td>
            <td>Novella</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate one">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Dead Souls</td>
            <td>Nikolai Gogol</td>
            <td>Picaresque</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="in-progress">
            <td>Lord of the Flies</td>
            <td>William Golding</td>
            <td>Allegorical Novel</td>
            <td><span class="status">In Progress</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Do Androids Dream of Electric Sheep?</td>
            <td>Philip K. Dick</td>
            <td>Science Fiction</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
        </tbody>
      </table>
    </div>
  </main>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
    font-family: Arial, sans-serif;
}

.table-container {
  overflow-x: auto;
}

table {
    border-collapse: collapse;
    width: 100vw;
}

thead {
    background-image: linear-gradient(hsl(197, 92%, 77%, 0.60), hsl(197, 92%, 50%, 0.60));
}

th {
    padding: 0.4em;
}

td {
    text-align: center;
    padding: 0.3em;
}

tr[class="read"] {
    background-image: linear-gradient(hsl(120, 100%, 85%, 0.6), hsl(120, 100%, 65%, 0.6));
}

tr[class="to-read"] {
    background-image: linear-gradient(hsl(120, 20%, 95%, 0.6), hsl(120, 20%, 75%, 0.6));
}

tr[class="in-progress"] {
    background-image: linear-gradient(hsl(40, 100%, 85%, 0.60), hsl(40, 100%, 65%, 0.6));
}

span {
    display: inline-block;
}

span[class="status"] {
    border-radius: 20%/60%;
}

span[class="status"],
span[class^="rate"] {
    height: 1.85em;
    width: 7em;
    padding: 0.3em;
}

tr[class="to-read"] span[class="status"] {
    border: 0.1em solid hsl(5, 100%, 45%);
    background-image: linear-gradient(hsl(5, 100%, 75%), hsl(5, 100%, 50%));
}

tr[class="read"] span[class="status"] {
    border: 0.1em solid hsl(120, 75%, 45%);
    background-image: linear-gradient(hsl(120, 75%, 75%), hsl(120, 75%, 50%));
}

tr[class="in-progress"] span[class="status"] {
    border: 0.1em solid hsl(40, 90%, 40%);
    background-image: linear-gradient(hsl(40, 90%, 75%), hsl(40, 90%, 50%));
}

span[class^="rate"] > span {
    border: 0.1em solid hsl(0, 0%, 50%);
    border-radius: 50%;
    background-color: hsl(0, 15%, 95%);
    height: 100%;
    width: 20%;
    margin: 0.1em;
}

span[class~="three"] span {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="two"] :nth-child(1),
span[class~="two"] :nth-child(2) {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="one"] :first-child {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}
```
