---
id: 69a013a86c97dac3784b2636
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

الآن، يجب أن تُظهر السجلات شيئًا مثل `{ head: { element: 42, next: null }, length: 1 }` في وحدة التحكم.

مع ذلك، هذه ليست الطريقة التي تتعامل بها مع حالة إضافة عقدة جديدة كأول عنصر في القائمة، أي الرأس.

للقيام بذلك بشكل صحيح، استخدم جملة if الشرطية للتحقق مما إذا كانت القائمة فارغة عن طريق استدعاء دالة `isEmpty` مع `list`. إذا تحقق الشرط، انقل سطر `list.head = node` إلى جسم جملة `if`.

# --hints--

يجب أن تحتوي على جملة `if` بشرط `isEmpty(list)`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*/);
```

يجب نقل السطر `list.head = node` إلى داخل جملة `if`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*list\.head\s*=\s*node;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

--fcc-editable-region--
    list.head = node;
--fcc-editable-region--
    list.length++;
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
console.log(myList);
console.log(isEmpty(myList));
```
