---
id: 69a016ac6edea7d8c8020c98
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

الآن، أنشئ حلقة `while` بشرط `current.next !== null`. هذا يعني أن الحلقة يجب أن تستمر طالما أن العقدة الحالية تملك عقدة أخرى بعدها.

# --hints--

يجب أن تنشئ حلقة `while` بشرط `current.next !== null`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'while(current.next!==null){');
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        --fcc-editable-region--

        --fcc-editable-region--
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
console.log(myList);
console.log(isEmpty(myList));
```
