---
id: 69a045fd236d0d973a59f103
title: الخطوة 19
challengeType: 1
dashedName: step-19
---

# --description--

إذا كان `previous` يساوي `null`، فهذا يعني أن العنصر المراد إزالته هو في بداية القائمة، لذا يجب تحديث `head` ليشير إلى العقدة التالية لإزالة العنصر الأول.

للقيام بذلك، أضف بند `else`. بداخله، عيّن `list.head` إلى `current.next`.

# --hints--

يجب أن تضيف بند `else` إلى جملة `if` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*/);
```

يجب أن تعيّن `list.head` إلى `current.next` داخل بند الـ `else`.

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*list\.head\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }

    if(previous !== null) {
        previous.next = current.next;
    --fcc-editable-region--
    } 
        
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
