---
id: 69a02ee82bbeaee00e44c67d
title: الخطوة 12
challengeType: 1
dashedName: step-12
---

# --description--

هذا كل ما تحتاجه لإضافة عنصر إلى القائمة. لرؤية الشكل النهائي، أضف رقمين آخرين من اختيارك إلى قائمة `myList`.

# --hints--

يجب أن تضيف رقمين آخرين من اختيارك إلى قائمة `myList`.

```js
const matches = __helpers.removeJSComments(code).match(/add\s*\(\s*myList\s*,\s*\d+\s*\);?/g) || [];

assert.isAtLeast(matches.length, 3);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
--fcc-editable-region--

--fcc-editable-region--
console.log(myList);
console.log(isEmpty(myList));
```
