---
id: 69a00a051b9a18694d00c16b
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستنفذ قائمة مرتبطة.

ابدأ بإنشاء دالة باسم `initList`. داخلها، أرجع كائنًا يحتوي على خاصية `head` بقيمة `null`، وخاصية `length` بقيمة `0`.

هذه هي الدالة التي ستستدعيها لتهيئة القائمة المرتبطة. تمثل `head` العنصر الأول في القائمة، وهي في البداية `null` لأنه لا يوجد عنصر في القائمة. تتعقب `length` عدد العقد (العناصر) في القائمة، والتي ستزيدها أو تنقصها عند إضافة أو إزالة عقدة من القائمة.

# --hints--

يجب أن تنشئ دالة `initList`.

```js
assert.isFunction(initList)
```

يجب أن تُرجع دالة `initList` كائنًا.

```js
assert.isObject(initList());
```

يجب أن تُرجع دالة `initList` كائنًا بخاصية `head` قيمتها `null`.

```js
assert.propertyVal(initList(), "head", null);
```

يجب أن تُرجع دالة `initList` كائنًا يحتوي على خاصية `length` بقيمة `0`.

```js
assert.propertyVal(initList(), "length", 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
