---
id: 69a03aaa1e6fe8152788ce6a
title: الخطوة 14
challengeType: 1
dashedName: step-14
---

# --description--

داخل دالة `remove`، استخدم `let` لتعريف متغير `previous` بقيمة `null`، واستخدم `let` مرة أخرى لتعريف متغير `current` بقيمة `list.head`. هذه المتغيرات ستستخدمها في الحلقة التي ستتصفح القائمة.

يتتبع `current` العقدة التي يتم فحصها حاليًا في القائمة، بدءًا من الرأس ويتقدم عقدة واحدة في كل مرة. أما `previous` فيتتبع العقدة التي تسبق العقدة الحالية مباشرة.

هذا مهم عند إزالة عقدة لأنك يجب أن تحدّث القيمة التالية للعقدة السابقة لتشير إلى العقدة التي تلي العقدة التي تُزال، للحفاظ على اتصال القائمة.

# --hints--

يجب أن تستخدم `let` لتعريف متغير `previous` بقيمة `null` داخل دالة `remove` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+previous\s*=\s*null;?/);
```

يجب أن تستخدم `let` لتعريف متغير `current` بقيمة `list.head` داخل دالة `remove` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
