---
id: 69a014bd9b1730e94136279f
title: الخطوة 8
challengeType: 1
dashedName: step-8
---

# --description--

إذا لم تكن القائمة فارغة، تريد إيجاد نهايتها، إضافة العنصر الجديد هناك، وجعله يشير إلى `null`. هذا يعني أنه يجب عليك المرور عبر القائمة بدءًا من الـ `head`.

لبدء ذلك، افتح بند `else`. بداخله، استخدم `let` لإنشاء متغير `current` واضبطه على `list.head`.

# --hints--

يجب أن يكون لديك بند `else`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{/);
```

يجب إنشاء متغير `current` مضبوط على `list.head` داخل بند الـ `else`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{\s*let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    --fcc-editable-region--
    } 
    --fcc-editable-region--

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
console.log(myList);
console.log(isEmpty(myList));
```
