---
id: 69a00cd4ec513345d40e95ec
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

الشيء التالي الذي ستعمل عليه هو دالة تتيح لك التحقق مما إذا كانت القائمة فارغة.

للقيام بذلك، أنشئ دالة باسم `isEmpty` تأخذ معلمة `list`. داخل الدالة، أرجع نتيجة تحقق مما إذا كانت الخاصية `length` في `list` تساوي تمامًا `0`.

# --hints--

يجب أن تنشئ دالة `isEmpty`.

```js
assert.isFunction(isEmpty)
```

يجب أن تحتوي دالة `isEmpty` على معلمة `list`.

```js
const regex = __helpers.functionRegex('isEmpty', ['list']);
assert.match(__helpers.removeJSComments(code), regex);
```

يجب أن تُرجع دالة `isEmpty` القيمة `list.length === 0`.

```js
const emptyList = { head: null, length: 0 };
const nonEmptyList = { head: {}, length: 1 };

assert.isTrue(isEmpty(emptyList));
assert.isFalse(isEmpty(nonEmptyList));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
