---
id: 69a01864d343dc2f2484d131
title: الخطوة 11
challengeType: 1
dashedName: step-11
---

# --description--

بعد العثور على العقدة الأخيرة، يجب ربط العقدة الجديدة بها. للقيام بذلك، عرّف `current.next` ليشير إلى `node` داخل حلقة while ولكن لا تزال ضمن بند الـ `else`.

# --hints--

يجب أن تعرّف `current.next` لتساوي `node` تحت حلقة `while`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\.next\s*=\s*node;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        --fcc-editable-region--

        --fcc-editable-region--
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
