---
id: 69a0179a85b00b6127e0af45
title: الخطوة 10
challengeType: 1
dashedName: step-10
---

# --description--

داخل حلقة `while`، حدّث `current` إلى `current.next` لكي تتحرك الحلقة للأمام عبر القائمة عقدة واحدة في كل مرة، وتتوقف عندما تصل إلى العقدة التي خاصية `next` فيها تساوي `null`، أي نهاية القائمة.

# --hints--

يجب تحديث قيمة المتغير `current` إلى `current.next` داخل حلقة `while`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            --fcc-editable-region--

            --fcc-editable-region--
        }
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
