---
id: 69a012c6224fac85832247dd
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

هل لاحظت نتائج `true`، `{ head: null, length: 0 }`، و`true` في وحدة التحكم؟ يحدث هذا لأنك لم تُعيّن أي شيء إلى القائمة `myList`، ولم تُزيد خاصية `length`.

تابع وقم بتعيين `node` إلى `list.head`، وتحت ذلك، زد خاصية `length` في `list`.

# --hints--

يجب تعيين `node` إلى `list.head`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'list.head=node');
```

يجب زيادة خاصية `length` في `list`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'list.length++');
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    --fcc-editable-region--

    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
console.log(myList);
console.log(isEmpty(myList));
```
