---
id: 69a030fbe487f185f55b8627
title: الخطوة 13
challengeType: 1
dashedName: step-13
---

# --description--

الآن، ستعمل على دالة لإزالة عقدة من القائمة.

أنشئ دالة `remove` بمعلمتين `list` و `element`.

# --hints--

يجب أن تنشئ دالة `remove`.

```js
assert.isFunction(remove)
```

يجب أن تحتوي دالة `remove` على المعلمتين `list` و `element`.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.remove.parameters;
assert.equal(parameters[0].toString(), "list")
assert.equal(parameters[1].toString(), "element")
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

--fcc-editable-region--

--fcc-editable-region--

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList);
console.log(isEmpty(myList));
```
