---
id: 69a00eaa43458e618331873c
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

الآن ستعمل على الدالة التي ستضيف عقدة إلى القائمة.

أنشئ دالة `add` بمعلمتين `list` و `element`.

# --hints--

يجب أن تنشئ دالة `add`.

```js
assert.isFunction(add)
```

يجب أن تحتوي دالة `add` على معلمتين `list` و `element`.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.add.parameters;
assert.equal(parameters[0].toString(), "list");
assert.equal(parameters[1].toString(), "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
