---
id: 69a04088e600bd8074a9852f
title: الخطوة 17
challengeType: 1
dashedName: step-17
---

# --description--

إذا كانت قيمة `current` تساوي `null`، فهذا يعني نهاية القائمة، لذا يجب أن تتوقف الدالة عن التنفيذ.

أنشئ عبارة `if` تتحقق مما إذا كانت قيمة `current` تساوي تمامًا `null`. داخل عبارة `if`، استخدم `return` للخروج من الدالة.

# --hints--

يجب أن تنشئ عبارة `if` بشرط `current === null`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(current\s*===\s*null\s*\)\s*\{\s*/);
```

يجب أن تستخدم الكلمة المفتاحية `return` داخل عبارة `if`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(current\s*===\s*null\s*\)\s*\{\s*return;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
