---
id: 69a00f673f171e29159e3ad8
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

تحتاج هذه العقدة إلى خاصيتين، العنصر الذي ستضيفه، ومؤشر إلى العقدة التالية.

داخل دالة add، أنشئ كائنًا باسم `node`، يجب أن يحتوي على خاصية باسم `element`، وخاصية `next` تُعيّن إلى `null`، وهو المؤشر إلى العنصر التالي في القائمة.

# --hints--

يجب أن تنشئ كائنًا باسم `node` داخل دالة `add` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(add.toString()), /const|var|let\s+node\s*=\s*{/);
```

يجب أن يحتوي كائن `node` الخاص بك على خاصية باسم `element` تُعيّن إلى `element`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /element\s*:\s*element/);
```

يجب أن يحتوي كائن `node` الخاص بك على خاصية باسم `next` تُعيّن إلى `null`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /next\s*:\s*null/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}
```
