---
id: 69a0102e761a245bb1cbd8e7
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

في هذه المرحلة، يجب أن تختبر القائمة المرتبطة لفهم أفضل.

عرّف متغيرًا باسم `myList` وضعه مساويًا لـ `initList()`. أسفله، سجّل استدعاءً للدالة `isEmpty()` باستخدام `myList` كمعلمة في وحدة التحكم. بعد ذلك، استدعِ الدالة `add` مع `myList` ورقم من اختيارك، ثم سجّل `myList` في وحدة التحكم. أسفل ذلك، سجّل استدعاءً آخر للدالة `isEmpty()` في وحدة التحكم باستخدام `myList` كمعلمة مرة أخرى.

# --hints--

يجب أن تعرّف متغيرًا باسم `myList` وتجعله مساويًا لـ `initList()`.

```js
assert.deepEqual(myList, initList());
```

يجب أن تسجل استدعاء `isEmpty(myList)` في وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(isEmpty\(myList\)\);?/)
```

يجب أن تستدعي الدالة `add` مع `myList` ورقم من اختيارك كمعلمات.

```js
assert.match(__helpers.removeJSComments(code), /add\s*\(\s*myList\s*,\s*\d+\s*\)/);
```

يجب أن تسجل متغير `myList` في وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?/)
```

يجب أن تسجل استدعاء `isEmpty(myList)` في وحدة التحكم مرة أخرى.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?\s*console\.log\(isEmpty\(myList\)\);?/)
```


# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
