---
id: 587d7b7d367417b2b2512b1d
title: for...in 문으로 객체의 키를 순차적으로 반복하기
challengeType: 1
forumTopicId: 301162
dashedName: iterate-through-the-keys-of-an-object-with-a-for---in-statement
---

# --description--

때때로 객체 내 모든 키를 순차적으로 반복해야 할 때가 있습니다. 이럴 때 <dfn>for...in</dfn> 반복문을 사용할 수 있습니다. for...in 반복문은 다음과 같습니다:

```javascript
const refrigerator = {
  'milk': 1,
  'eggs': 12,
};

for (const food in refrigerator) {
  console.log(food, refrigerator[food]);
}
```

이 코드는 `milk 1` 와 `eggs 12` 을 각각의 키-값 쌍을 한 줄씩 출력합니다.

반복문 헤드에서 변수 `food` 를 정의했고, 이 변수는 반복할 때마다 객체의 각 키로 설정되어 각 음식 이름이 콘솔에 출력됩니다.

**참고:** 객체는 배열처럼 저장된 키의 순서를 유지하지 않으므로, 객체 내 키의 위치나 나타나는 상대적 순서는 해당 키를 참조하거나 접근할 때 중요하지 않습니다.

# --instructions--

하나의 인수 `countOnline` 를 받는 함수 `allUsers` 를 정의했습니다. 이 함수 안에서 <dfn>for...in</dfn> 문을 사용해 `allUsers` 객체의 키를 순차적으로 반복하고, `online` 속성이 `true` 으로 설정된 사용자의 수를 반환하세요. `countOnline` 에 전달될 수 있는 객체 예시는 아래에 나와 있습니다. 각 사용자는 `online` 속성이 `true` 또는 `false` 로 설정되어 있습니다.

```js
{
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}
```

# --hints--

함수 `countOnline` 는 전달받은 객체의 키를 반복하기 위해 `for in` 문을 사용해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(code),
    /for\s*\(\s*(var|let|const)\s+[a-zA-Z_$]\w*\s+in\s+[a-zA-Z_$]\w*\s*\)/
);
```

함수 `countOnline` 는 객체 `1` 이 전달될 때 `{ Alan: { online: false }, Jeff: { online: true }, Sarah: { online: false } }` 를 반환해야 합니다.

```js
const usersObj1 = {
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}
assert.equal(countOnline(usersObj1), 1);
```

함수 `countOnline` 는 객체 `2` 이 전달될 때 `{ Alan: { online: true }, Jeff: { online: false }, Sarah: { online: true } }` 를 반환해야 합니다.

```js
const usersObj2 = {
  Alan: {
    online: true
  },
  Jeff: {
    online: false
  },
  Sarah: {
    online: true
  }
}
assert.equal(countOnline(usersObj2), 2);
```

함수 `countOnline` 는 객체 `0` 가 전달될 때 `{ Alan: { online: false }, Jeff: { online: false }, Sarah: { online: false } }` 를 반환해야 합니다.

```js
const usersObj3 = {
  Alan: {
    online: false
  },
  Jeff: {
    online: false
  },
  Sarah: {
    online: false
  }
}
assert.equal(countOnline(usersObj3), 0);
```

# --seed--

## --seed-contents--

```js
const users = {
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}

function countOnline(allUsers) {
  // Only change code below this line

  // Only change code above this line
}

console.log(countOnline(users));
```

# --solutions--

```js
function countOnline(allUsers) {
  let numOnline = 0;
  for(const user in allUsers){
    if(allUsers[user].online) {
      numOnline++;
    }
  }
  return numOnline;
}
```
