---
id: 587d7b7b367417b2b2512b13
title: 전개 연산자로 배열 복사하기
challengeType: 1
forumTopicId: 301157
dashedName: copy-an-array-with-the-spread-operator
---

# --description--

`slice()`는 배열의 어떤 요소를 선택적으로 복사할 수 있게 해주고, 여러 유용한 작업 중 하나지만, ES6의 새로운 <dfn>전개 연산자</dfn>는 간단하고 매우 읽기 쉬운 구문으로 배열의 *모든* 요소를 순서대로 쉽게 복사할 수 있게 해줍니다. 전개 구문은 단순히 이렇게 생겼습니다: `...`

실제로 전개 연산자를 사용해 배열을 이렇게 복사할 수 있습니다:

```js
let thisArray = [true, true, undefined, false, null];
let thatArray = [...thisArray];
```

`thatArray`는 `[true, true, undefined, false, null]`와 같습니다. `thisArray`는 변경되지 않고 `thatArray`는 `thisArray`와 같은 요소를 포함합니다.

# --instructions--

`copyMachine`라는 함수를 정의했으며, 이 함수는 `arr` (배열)과 `num` (숫자)를 인수로 받습니다. 이 함수는 `num`의 `arr`개 복사본으로 이루어진 새 배열을 반환해야 합니다. 대부분의 작업은 완료했지만 아직 제대로 작동하지 않습니다. 전개 구문을 사용해 함수를 수정하여 올바르게 작동하게 하세요(힌트: 이미 배운 다른 메서드가 도움이 될 수 있습니다!).

# --hints--

`copyMachine([true, false, true], 2)`는 `[[true, false, true], [true, false, true]]`를 반환해야 합니다

```js
assert.deepEqual(copyMachine([true, false, true], 2), [
  [true, false, true],
  [true, false, true]
]);
```

`copyMachine([1, 2, 3], 5)`는 `[[1, 2, 3], [1, 2, 3], [1, 2, 3], [1, 2, 3], [1, 2, 3]]`를 반환해야 합니다

```js
assert.deepEqual(copyMachine([1, 2, 3], 5), [
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3]
]);
```

`copyMachine([true, true, null], 1)`는 `[[true, true, null]]`를 반환해야 합니다

```js
assert.deepEqual(copyMachine([true, true, null], 1), [[true, true, null]]);
```

`copyMachine(["it works"], 3)`는 `[["it works"], ["it works"], ["it works"]]`를 반환해야 합니다

```js
assert.deepEqual(copyMachine(['it works'], 3), [
  ['it works'],
  ['it works'],
  ['it works']
]);
```

`copyMachine` 함수는 `spread operator` 배열에 `arr`를 활용해야 합니다

```js
assert(__helpers.removeJSComments(__helpers.removeJSComments(code)).match(/\.\.\.\s*arr/));
```

# --seed--

## --seed-contents--

```js
function copyMachine(arr, num) {
  let newArr = [];
  while (num >= 1) {
    // Only change code below this line

    // Only change code above this line
    num--;
  }
  return newArr;
}

console.log(copyMachine([true, false, true], 2));
```

# --solutions--

```js
function copyMachine(arr,num){
    let newArr=[];
    while(num >=1){
    newArr.push([...arr]);
    num--;
    }
    return newArr;
}
console.log(copyMachine([true, false, true], 2));
```
