---
id: 587d7b7c367417b2b2512b18
title: JavaScript 객체에 키-값 쌍 추가하기
challengeType: 1
forumTopicId: 301153
dashedName: add-key-value-pairs-to-javascript-objects
---

# --description--

객체는 가장 기본적으로 키-값 쌍의 모음입니다. 다시 말해, 고유 식별자인 속성(키)에 매핑된 데이터(값) 조각입니다. 예제를 살펴보세요:

```js
const tekkenCharacter = {
  player: 'Hwoarang',
  fightingStyle: 'Tae Kwon Doe',
  human: true
};
```

위 코드는 `tekkenCharacter`라는 Tekken 비디오 게임 캐릭터 객체를 정의합니다. 세 개의 속성이 있으며 각각 특정 값에 매핑되어 있습니다. "origin" 같은 추가 속성을 넣고 싶다면, 객체에 `origin`을 할당하여 할 수 있습니다:

```js
tekkenCharacter.origin = 'South Korea';
```

이것은 점 표기법을 사용합니다. `tekkenCharacter` 객체를 보면 이제 `origin` 속성이 포함되어 있을 것입니다. Hwoarang은 또한 독특한 주황색 머리카락을 가지고 있었습니다. 이 속성은 대괄호 표기법으로 다음과 같이 추가할 수 있습니다:

```js
tekkenCharacter['hair color'] = 'dyed orange';
```

속성에 공백이 있거나 속성 이름을 변수로 사용하고 싶을 때는 대괄호 표기법이 필요합니다. 위 예제에서는 속성이 문자열임을 나타내기 위해 따옴표로 감싸져 있으며, 그대로 추가됩니다. 따옴표가 없으면 변수로 평가되어 속성 이름은 변수 값이 됩니다. 변수 예제는 다음과 같습니다:

```js
const eyes = 'eye color';

tekkenCharacter[eyes] = 'brown';
```

모든 예제를 추가한 후 객체는 다음과 같이 보일 것입니다:

```js
{
  player: 'Hwoarang',
  fightingStyle: 'Tae Kwon Doe',
  human: true,
  origin: 'South Korea',
  'hair color': 'dyed orange',
  'eye color': 'brown'
};
```

# --instructions--

`foods` 객체가 세 개의 항목과 함께 생성되었습니다. 원하는 구문을 사용해 `bananas`에 `13` 값을, `grapes`에 `35` 값을, `strawberries`에 `27` 값을 각각 추가하세요.

# --hints--

`foods`은 객체여야 합니다.

```js
assert(typeof foods === 'object');
```

`foods` 객체는 `bananas` 키에 `13` 값을 가져야 합니다.

```js
assert(foods.bananas === 13);
```

`foods` 객체는 `grapes` 키에 `35` 값을 가져야 합니다.

```js
assert(foods.grapes === 35);
```

`foods` 객체는 `strawberries` 키에 `27` 값을 가져야 합니다.

```js
assert(foods.strawberries === 27);
```

`foods` 객체 정의는 변경하지 마세요.

```js
assert(
  __helpers.removeJSComments(code).search(/let foods/) === -1 &&
  __helpers.removeJSComments(code).search(/const\s+foods\s*=\s*{\s*apples:\s*25,\s*oranges:\s*32,\s*plums:\s*28\s*};/
) !== -1
);
```

# --seed--

## --seed-contents--

```js
const foods = {
  apples: 25,
  oranges: 32,
  plums: 28
};

// Only change code below this line

// Only change code above this line

console.log(foods);
```

# --solutions--

```js
const foods = {
  apples: 25,
  oranges: 32,
  plums: 28
};

foods['bananas'] = 13;
foods['grapes']  = 35;
foods['strawberries'] = 27;
```
