---
id: 587d7b7a367417b2b2512b12
title: slice()를 사용해 배열 항목 복사하기
challengeType: 1
forumTopicId: 301158
dashedName: copy-array-items-using-slice
---

# --description--

다음으로 다룰 메서드는 `slice()`입니다. 배열을 수정하는 대신 `slice()`은 주어진 개수의 요소를 새 배열로 *복사*하거나 *추출*하여 호출된 배열은 그대로 둡니다. `slice()`는 매개변수를 2개만 받는데 — 첫 번째는 추출을 시작할 인덱스이고, 두 번째는 추출을 멈출 인덱스입니다(이 인덱스에 해당하는 요소는 포함하지 않고 그 전까지 추출합니다). 다음을 살펴보세요:

```js
let weatherConditions = ['rain', 'snow', 'sleet', 'hail', 'clear'];

let todaysWeather = weatherConditions.slice(1, 3);
```

`todaysWeather`는 `['snow', 'sleet']` 값을 가지게 되고, `weatherConditions`는 여전히 `['rain', 'snow', 'sleet', 'hail', 'clear']` 값을 유지합니다.

결과적으로 기존 배열에서 요소를 추출해 새 배열을 만든 셈입니다.

# --instructions--

배열을 인수로 받는 `forecast` 함수를 정의했습니다. `slice()`을 사용해 인수 배열에서 정보를 추출하고 문자열 요소 `warm`와 `sunny`를 포함하는 새 배열을 반환하도록 함수를 수정하세요.

# --hints--

`forecast`는 `["warm", "sunny"]`를 반환해야 합니다.

```js
assert.deepEqual(
  forecast(['cold', 'rainy', 'warm', 'sunny', 'cool', 'thunderstorms']),
  ['warm', 'sunny']
);
```

`forecast` 함수는 `slice()` 메서드를 활용해야 합니다.

```js
assert(/\.slice\(/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function forecast(arr) {
  // Only change code below this line

  return arr;
}

// Only change code above this line
console.log(forecast(['cold', 'rainy', 'warm', 'sunny', 'cool', 'thunderstorms']));
```

# --solutions--

```js
function forecast(arr) {
  return arr.slice(2,4);
}
```
