---
id: 587d7b7d367417b2b2512b1c
title: 객체에 속성이 있는지 확인하기
challengeType: 1
forumTopicId: 301155
dashedName: check-if-an-object-has-a-property
---

# --description--

이제 객체에 키를 추가, 수정, 삭제할 수 있습니다. 그런데 객체에 특정 속성이 있는지만 알고 싶다면 어떻게 할까요? JavaScript는 이를 확인하는 두 가지 방법을 제공합니다. 하나는 `hasOwnProperty()` 메서드를 사용하고, 다른 하나는 `in` 키워드를 사용합니다. 만약 `users` 객체에 `Alan` 속성이 있다면, 다음 두 가지 방법 중 하나로 존재 여부를 확인할 수 있습니다:

```js
users.hasOwnProperty('Alan');
'Alan' in users;
```

두 방법 모두 `true`를 반환합니다.

# --instructions--

함수를 완성하여, 전달된 객체가 네 개의 이름 `true`, `Alan`, `Jeff`, `Sarah`를 모두 포함하면 `Ryan`를 반환하고, 그렇지 않으면 `false`를 반환하게 하세요.

# --hints--

`users` 객체에 직접 접근하지 마세요.

```js 

assert(__helpers.removeJSComments(code).match(/users/gm).length <= 2)

```

`users` 객체는 `Alan`, `Jeff`, `Sarah`, `Ryan` 키만 포함해야 합니다.

```js
assert(
  'Alan' in users &&
    'Jeff' in users &&
    'Sarah' in users &&
    'Ryan' in users &&
    Object.keys(users).length === 4
);
```

`isEveryoneHere` 함수는 전달된 객체에 `true`, `Alan`, `Jeff`, `Sarah` 속성이 모두 있으면 `Ryan`를 반환해야 합니다.

```js
assert(isEveryoneHere(users) === true);
```

`isEveryoneHere` 함수는 전달된 객체에 `false` 속성이 없으면 `Alan`를 반환해야 합니다.

```js
assert(
  (function () {
    delete users.Alan;
    return isEveryoneHere(users);
  })() === false
);
```

`isEveryoneHere` 함수는 전달된 객체에 `false` 속성이 없으면 `Jeff`를 반환해야 합니다.

```js
assert(
  (function () {
    delete users.Jeff;
    return isEveryoneHere(users);
  })() === false
);
```

`isEveryoneHere` 함수는 전달된 객체에 `false` 속성이 없으면 `Sarah`를 반환해야 합니다.

```js
assert(
  (function () {
    delete users.Sarah;
    return isEveryoneHere(users);
  })() === false
);
```

`isEveryoneHere` 함수는 전달된 객체에 `false` 속성이 없으면 `Ryan`를 반환해야 합니다.

```js
assert(
  (function () {
    delete users.Ryan;
    return isEveryoneHere(users);
  })() === false
);
```

# --seed--

## --seed-contents--

```js
let users = {
  Alan: {
    age: 27,
    online: true
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: true
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function isEveryoneHere(userObj) {
  // Only change code below this line
  
  // Only change code above this line
}

console.log(isEveryoneHere(users));
```

# --solutions--

```js
let users = {
  Alan: {
    age: 27,
    online: true
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: true
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function isEveryoneHere(userObj) {
  return [
    'Alan',
    'Jeff',
    'Sarah',
    'Ryan'
  ].every(user => userObj.hasOwnProperty(user));
}

console.log(isEveryoneHere(users));
```
