---
id: 587d78b2367417b2b2512b0f
title: pop()와 shift()로 배열에서 항목 제거하기
challengeType: 1
forumTopicId: 301165
dashedName: remove-items-from-an-array-with-pop-and-shift
---

# --description--

`push()`와 `unshift()`에는 거의 기능적으로 반대되는 메서드가 있습니다: `pop()`와 `shift()`. 이제 짐작했겠지만, 추가하는 대신 `pop()`는 배열 끝에서 요소를 *제거*하고, `shift()`는 배열 시작에서 요소를 제거합니다. `pop()`와 `shift()` 그리고 그들의 대응 메서드 `push()`와 `unshift()`의 주요 차이점은 두 메서드 모두 매개변수를 받지 않고, 한 번에 배열에서 단일 요소만 수정할 수 있다는 점입니다.

한번 살펴보겠습니다:

```js
let greetings = ['whats up?', 'hello', 'see ya!'];

greetings.pop();
```

`greetings`는 `['whats up?', 'hello']` 값을 갖게 됩니다.

```js
greetings.shift();
```

`greetings`는 `['hello']` 값을 갖게 됩니다.

또한 두 메서드 모두 제거된 요소의 값을 이렇게 반환할 수 있습니다:

```js
let popped = greetings.pop();
```

`greetings`는 `[]` 값을 갖고, `popped`는 `hello` 값을 갖게 됩니다.

# --instructions--

배열을 인수로 받고 새 배열을 반환하는 `popShift` 함수를 정의했습니다. `pop()`와 `shift()`를 사용해 인수 배열의 첫 번째와 마지막 요소를 제거하고, 제거된 요소들을 각각 대응하는 변수에 할당하도록 함수를 수정하세요. 반환되는 배열에는 이 변수들의 값이 포함되어야 합니다.

# --hints--

`popShift(["challenge", "is", "not", "complete"])`는 `["challenge", "complete"]`를 반환해야 합니다.

```js
assert.deepEqual(popShift(['challenge', 'is', 'not', 'complete']), [
  'challenge',
  'complete'
]);
```

`popShift` 함수는 `pop()` 메서드를 활용해야 합니다.

```js
assert.notStrictEqual(popShift.toString().search(/\.pop\(/), -1);
```

`popShift` 함수는 `shift()` 메서드를 활용해야 합니다.

```js
assert.notStrictEqual(popShift.toString().search(/\.shift\(/), -1);
```

# --seed--

## --seed-contents--

```js
function popShift(arr) {
  let popped; // Change this line
  let shifted; // Change this line
  return [shifted, popped];
}

console.log(popShift(['challenge', 'is', 'not', 'complete']));
```

# --solutions--

```js
function popShift(arr) {
  let popped = arr.pop(); // Change this line
  let shifted = arr.shift(); // Change this line
  return [shifted, popped];
}
```
