---
id: 587d7b7c367417b2b2512b19
title: 객체 안에 중첩된 객체 수정하기
challengeType: 1
forumTopicId: 301164
dashedName: modify-an-object-nested-within-an-object
---

# --description--

이제 조금 더 복잡한 객체를 살펴보겠습니다. 객체 속성은 임의의 깊이로 중첩될 수 있으며, 그 값은 배열이나 다른 객체를 포함해 JavaScript가 지원하는 모든 데이터 유형일 수 있습니다. 다음을 참고하세요:

```js
let nestedObject = {
  id: 28802695164,
  date: 'December 31, 2016',
  data: {
    totalUsers: 99,
    online: 80,
    onlineStatus: {
      active: 67,
      away: 13,
      busy: 8
    }
  }
};
```

`nestedObject`에는 세 개의 속성이 있습니다: `id` (값은 숫자), `date` (값은 문자열), 그리고 `data` (중첩된 구조를 가진 객체 값). 구조가 빠르게 복잡해질 수 있지만, 필요한 정보를 접근하는 데는 여전히 같은 표기법을 사용할 수 있습니다. 중첩된 `10` 객체의 `busy` 속성에 `onlineStatus` 값을 할당하려면 점 표기법을 사용해 속성을 참조합니다:

```js
nestedObject.data.onlineStatus.busy = 10;
```

# --instructions--

여기 `userActivity` 객체를 정의했으며, 그 안에 또 다른 객체가 중첩되어 있습니다. `online` 키의 값을 `45`로 설정하세요.

# --hints--

`userActivity`에는 `id`, `date`, `data` 속성이 있어야 합니다.

```js
assert(
  'id' in userActivity && 'date' in userActivity && 'data' in userActivity
);
```

`userActivity`에는 `data` 키가 있고, 그 값은 `totalUsers`과 `online` 키를 가진 객체여야 합니다.

```js
assert('totalUsers' in userActivity.data && 'online' in userActivity.data);
```

`online`의 `data` 키에 중첩된 `userActivity` 속성은 `45`로 설정해야 합니다.

```js
assert(userActivity.data.online === 45);
```

`online` 속성은 점 표기법이나 대괄호 표기법을 사용해 설정하세요.

```js
assert.strictEqual(__helpers.removeJSComments(code).search(/online: 45/), -1);
```

# --seed--

## --seed-contents--

```js
let userActivity = {
  id: 23894201352,
  date: 'January 1, 2017',
  data: {
    totalUsers: 51,
    online: 42
  }
};

// Only change code below this line

// Only change code above this line

console.log(userActivity);
```

# --solutions--

```js
let userActivity = {
  id: 23894201352,
  date: 'January 1, 2017',
  data: {
    totalUsers: 51,
    online: 42
  }
};

userActivity.data.online = 45;
```
