---
id: 587d7b7b367417b2b2512b17
title: 전개 연산자로 배열 합치기
challengeType: 1
forumTopicId: 301156
dashedName: combine-arrays-with-the-spread-operator
---

# --description--

<dfn>전개</dfn> 연산자의 또 다른 큰 장점은 배열을 합치거나 한 배열의 모든 요소를 다른 배열의 원하는 인덱스에 삽입할 수 있다는 점입니다. 기존 구문으로는 배열을 연결할 수 있지만, 이는 한 배열의 끝과 다른 배열의 시작 부분에서만 합칠 수 있습니다. 전개 구문을 사용하면 다음 작업이 매우 간단해집니다:

```js
let thisArray = ['sage', 'rosemary', 'parsley', 'thyme'];

let thatArray = ['basil', 'cilantro', ...thisArray, 'coriander'];
```

`thatArray`는 `['basil', 'cilantro', 'sage', 'rosemary', 'parsley', 'thyme', 'coriander']` 값을 갖게 됩니다.

전개 구문을 사용해 전통적인 방법보다 더 복잡하고 장황했을 작업을 방금 수행했습니다.

# --instructions--

`spreadOut`라는 함수를 정의했고, 이 함수는 변수 `sentence`를 반환합니다. <dfn>전개</dfn> 연산자를 사용해 이 함수를 수정하여 배열 `['learning', 'to', 'code', 'is', 'fun']`를 반환하게 만드세요.

# --hints--

`spreadOut`는 `["learning", "to", "code", "is", "fun"]`를 반환해야 합니다.

```js
assert.deepEqual(spreadOut(), ['learning', 'to', 'code', 'is', 'fun']);
```

`spreadOut` 함수는 전개 구문을 활용해야 합니다.

```js
assert.notStrictEqual(spreadOut.toString().search(/[...]/), -1);
```

# --seed--

## --seed-contents--

```js
function spreadOut() {
  let fragment = ['to', 'code'];
  let sentence; // Change this line
  return sentence;
}

console.log(spreadOut());
```

# --solutions--

```js
function spreadOut() {
  let fragment = ['to', 'code'];
  let sentence = ['learning', ...fragment, 'is', 'fun'];
  return sentence;
}
```
