---
id: 587d7b7b367417b2b2512b15
title: 배열의 모든 항목을 for 루프로 순차적으로 반복하기
challengeType: 1
forumTopicId: 301161
dashedName: iterate-through-all-an-arrays-items-using-for-loops
---

# --description--

배열을 다룰 때, 필요한 하나 이상의 요소를 찾거나 특정 조건을 만족하는 데이터 항목을 기준으로 배열을 조작하기 위해 각 항목을 순차적으로 반복하는 것이 매우 유용할 때가 있습니다. JavaScript는 각각 약간씩 다른 방식으로 배열을 순회하여 다양한 결과를 얻는 여러 내장 메서드를 제공합니다(`every()`, `forEach()`, `map()` 등), 하지만 가장 유연하고 제어가 가장 많은 기법은 간단한 `for` 루프입니다.

다음 예제를 살펴보세요:

```js
function greaterThanTen(arr) {
  let newArr = [];
  for (let i = 0; i < arr.length; i++) {
    if (arr[i] > 10) {
      newArr.push(arr[i]);
    }
  }
  return newArr;
}

greaterThanTen([2, 12, 8, 14, 80, 0, 1]);
```

`for` 루프를 사용하여 이 함수는 배열의 각 요소를 순차적으로 반복하며 접근하고, 우리가 만든 간단한 테스트를 적용합니다. 이렇게 해서 `10`보다 큰 데이터 항목을 쉽게 프로그래밍적으로 판별하고, 그 항목들을 포함하는 새로운 배열 `[12, 14, 80]`를 반환했습니다.

# --instructions--

`filteredArray`라는 함수를 정의했습니다. 이 함수는 중첩된 배열인 `arr`과 `elem`를 인수로 받아 새로운 배열을 반환합니다. `elem`은 `arr` 내 중첩된 하나 이상의 배열에 있을 수도 있고 없을 수도 있는 요소를 나타냅니다. `for` 루프를 사용하여 전달된 배열의 필터링된 버전을 반환하도록 함수를 수정하세요. 이때 `arr` 내 중첩된 배열 중 `elem`를 포함하는 배열은 제거되어야 합니다.

# --hints--

`filteredArray([[10, 8, 3], [14, 6, 23], [3, 18, 6]], 18)`는 `[[10, 8, 3], [14, 6, 23]]`를 반환해야 합니다.

```js
assert.deepEqual(
  filteredArray(
    [
      [10, 8, 3],
      [14, 6, 23],
      [3, 18, 6]
    ],
    18
  ),
  [
    [10, 8, 3],
    [14, 6, 23]
  ]
);
```

`filteredArray([["trumpets", 2], ["flutes", 4], ["saxophones", 2]], 2)`는 `[["flutes", 4]]`를 반환해야 합니다.

```js
assert.deepEqual(
  filteredArray(
    [
      ['trumpets', 2],
      ['flutes', 4],
      ['saxophones', 2]
    ],
    2
  ),
  [['flutes', 4]]
);
```

`filteredArray([["amy", "beth", "sam"], ["dave", "sean", "peter"]], "peter")`는 `[["amy", "beth", "sam"]]`를 반환해야 합니다.

```js
assert.deepEqual(
  filteredArray(
    [
      ['amy', 'beth', 'sam'],
      ['dave', 'sean', 'peter']
    ],
    'peter'
  ),
  [['amy', 'beth', 'sam']]
);
```

`filteredArray([[3, 2, 3], [1, 6, 3], [3, 13, 26], [19, 3, 9]], 3)`는 `[]`를 반환해야 합니다.

```js
assert.deepEqual(
  filteredArray(
    [
      [3, 2, 3],
      [1, 6, 3],
      [3, 13, 26],
      [19, 3, 9]
    ],
    3
  ),
  []
);
```

`filteredArray` 함수는 `for` 루프를 활용해야 합니다.

```js
assert.notStrictEqual(filteredArray.toString().search(/for/), -1);
```

# --seed--

## --seed-contents--

```js
function filteredArray(arr, elem) {
  let newArr = [];
  // Only change code below this line

  // Only change code above this line
  return newArr;
}

console.log(filteredArray([[3, 2, 3], [1, 6, 3], [3, 13, 26], [19, 3, 9]], 3));
```

# --solutions--

```js
function filteredArray(arr, elem) {
  let newArr = [];
  for (let i = 0; i<arr.length; i++) {
    if (arr[i].indexOf(elem) < 0) {
      newArr.push(arr[i]);
    }
  }
  return newArr;
}
```
