---
id: 587d7b7b367417b2b2512b16
title: 복잡한 다차원 배열 만들기
challengeType: 1
forumTopicId: 301159
dashedName: create-complex-multi-dimensional-arrays
---

# --description--

멋져요! 배열에 대해 정말 많은 것을 배웠습니다! 지금까지는 꽤 높은 수준의 개요였고, 배열을 다루는 데에는 앞으로 더 많은 것을 배우게 될 것입니다. 이후 섹션에서 많이 보게 될 거예요. 하지만 <dfn>객체</dfn>를 살펴보기 전에, 배열이 이전 도전 과제에서 본 것보다 어떻게 더 복잡해질 수 있는지 한 번 더 살펴보겠습니다.

배열을 자료 구조로 생각할 때 가장 강력한 기능 중 하나는 배열이 다른 배열을 포함하거나 심지어 완전히 다른 배열로 구성될 수 있다는 점입니다. 이전 도전 과제에서 배열 안에 배열이 포함된 것을 보았지만, 꽤 단순한 형태였습니다. 그러나 배열은 무한한 깊이로 다른 배열을 포함할 수 있고, 각각 임의의 깊이를 가질 수 있습니다. 이렇게 배열은 매우 빠르게 <dfn>다차원</dfn> 또는 중첩된 배열이라는 매우 복잡한 자료 구조가 될 수 있습니다. 다음 예제를 보세요:

```js
let nestedArray = [
  ['deep'],
  [
    ['deeper'], ['deeper'] 
  ],
  [
    [
      ['deepest'], ['deepest']
    ],
    [
      [
        ['deepest-est?']
      ]
    ]
  ]
];
```

`deep` 배열은 2단계 깊이로 중첩되어 있습니다. `deeper` 배열은 3단계 깊이입니다. `deepest` 배열은 4단계, `deepest-est?`은 5단계입니다.

이 예제가 복잡해 보일 수 있지만, 대량의 데이터를 다룰 때 이런 수준의 복잡성은 드물지 않고, 심지어 흔한 편입니다. 하지만 대괄호 표기법을 사용하면 이렇게 복잡한 배열의 가장 깊은 단계도 아주 쉽게 접근할 수 있습니다:

```js
console.log(nestedArray[2][1][0][0][0]);
```

이 코드는 `deepest-est?` 문자열을 출력합니다. 그리고 이제 그 데이터가 어디에 있는지 알았으니, 필요하다면 값을 다시 설정할 수도 있습니다:

```js
nestedArray[2][1][0][0][0] = 'deeper still';

console.log(nestedArray[2][1][0][0][0]);
```

이제 `deeper still`를 출력합니다.

# --instructions--

`myNestedArray`라는 변수를 배열로 정의했습니다. `myNestedArray`을 수정하여 <dfn>문자열</dfn>, <dfn>숫자</dfn>, <dfn>불리언</dfn>을 데이터 요소로 조합해 정확히 5단계 깊이를 갖도록 만드세요(가장 바깥 배열이 1단계임을 기억하세요). 3단계 중간 어딘가에 `deep` 문자열을, 4단계에는 `deeper` 문자열을, 5단계에는 `deepest` 문자열을 포함하세요.

# --hints--

`myNestedArray`에는 숫자, 불리언, 문자열만 데이터 요소로 포함되어야 합니다

```js
assert.strictEqual(
  (function (arr) {
    let flattened = (function flatten(arr) {
      const flat = [].concat(...arr);
      return flat.some(Array.isArray) ? flatten(flat) : flat;
    })(arr);
    for (let i = 0; i < flattened.length; i++) {
      if (
        typeof flattened[i] !== 'number' &&
        typeof flattened[i] !== 'string' &&
        typeof flattened[i] !== 'boolean'
      ) {
        return false;
      }
    }
    return true;
  })(myNestedArray),
  true
);
```

`myNestedArray`은 정확히 5단계 깊이를 가져야 합니다

```js
assert.strictEqual(
  (function (arr) {
    let depth = 0;
    function arrayDepth(array, i, d) {
      if (Array.isArray(array[i])) {
        arrayDepth(array[i], 0, d + 1);
      } else {
        depth = d > depth ? d : depth;
      }
      if (i < array.length) {
        arrayDepth(array, i + 1, d);
      }
    }
    arrayDepth(arr, 0, 0);
    return depth;
  })(myNestedArray),
  4
);
```

`myNestedArray`에는 3단계 깊이로 중첩된 배열에 `deep` 문자열이 정확히 한 번 포함되어야 합니다

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deep').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deep')[0] === 2
);
```

`myNestedArray`에는 4단계 깊이로 중첩된 배열에 `deeper` 문자열이 정확히 한 번 포함되어야 합니다

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deeper').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deeper')[0] === 3
);
```

`myNestedArray`에는 5단계 깊이로 중첩된 배열에 `deepest` 문자열이 정확히 한 번 포함되어야 합니다

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deepest').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deepest')[0] === 4
);
```

# --seed--

## --seed-contents--

```js
let myNestedArray = [
  // Only change code below this line
  ['unshift', false, 1, 2, 3, 'complex', 'nested'],
  ['loop', 'shift', 6, 7, 1000, 'method'],
  ['concat', false, true, 'spread', 'array'],
  ['mutate', 1327.98, 'splice', 'slice', 'push'],
  ['iterate', 1.3849, 7, '8.4876', 'arbitrary', 'depth']
  // Only change code above this line
];
```

# --solutions--

```js
let myNestedArray = [
  ['unshift', ['deep', ['deeper', ['deepest']]],false, 1, 2, 3, 'complex', 'nested'],
  ['loop', 'shift', 6, 7, 1000, 'method'],
  ['concat', false, true, 'spread', 'array'],
  ['mutate', 1327.98, 'splice', 'slice', 'push'],
  ['iterate', 1.3849, 7, '8.4876', 'arbitrary', 'depth']
];
```
