---
id: 587d78b2367417b2b2512b0e
title: push()와 unshift()로 배열에 항목 추가하기
challengeType: 1
forumTopicId: 301151
dashedName: add-items-to-an-array-with-push-and-unshift
---

# --description--

배열의 길이와 배열이 포함할 수 있는 데이터 유형은 고정되어 있지 않습니다. 배열은 원하는 길이만큼 정의할 수 있고, 시간이 지나면서 요소를 추가하거나 제거할 수 있습니다. 즉, 배열은 <dfn>변경 가능</dfn>합니다. 이번 도전 과제에서는 배열을 프로그래밍적으로 수정할 수 있는 두 가지 메서드인 `Array.push()`와 `Array.unshift()`를 살펴보겠습니다.

두 메서드는 하나 이상의 요소를 매개변수로 받아 메서드가 호출된 배열에 그 요소들을 추가합니다. `push()` 메서드는 배열의 끝에 요소를 추가하고, `unshift()`는 배열의 시작 부분에 요소를 추가합니다. 다음을 참고하세요:

```js
let twentyThree = 'XXIII';
let romanNumerals = ['XXI', 'XXII'];

romanNumerals.unshift('XIX', 'XX');
```

`romanNumerals`는 `['XIX', 'XX', 'XXI', 'XXII']` 값을 갖게 됩니다.

```js
romanNumerals.push(twentyThree);
```

`romanNumerals`는 `['XIX', 'XX', 'XXI', 'XXII', 'XXIII']` 값을 갖게 됩니다. 변수도 전달할 수 있다는 점에 주목하세요. 이는 배열 데이터를 동적으로 수정할 때 훨씬 더 큰 유연성을 제공합니다.

# --instructions--

배열을 인수로 전달받는 `mixedNumbers` 함수를 정의했습니다. 이 함수를 수정하여 `push()`와 `unshift()`를 사용해 배열의 시작 부분에 `'I', 2, 'three'`를, 끝 부분에 `7, 'VIII', 9`를 추가하세요. 반환된 배열은 1부터 9까지 숫자를 순서대로 나타내야 합니다.

# --hints--

`mixedNumbers(["IV", 5, "six"])`는 이제 `["I", 2, "three", "IV", 5, "six", 7, "VIII", 9]`를 반환해야 합니다.

```js
assert.deepEqual(mixedNumbers(['IV', 5, 'six']), [
  'I',
  2,
  'three',
  'IV',
  5,
  'six',
  7,
  'VIII',
  9
]);
```

`mixedNumbers` 함수는 `push()` 메서드를 활용해야 합니다.

```js
assert(mixedNumbers.toString().match(/\.push/));
```

`mixedNumbers` 함수는 `unshift()` 메서드를 활용해야 합니다.

```js
assert(mixedNumbers.toString().match(/\.unshift/));
```

# --seed--

## --seed-contents--

```js
function mixedNumbers(arr) {
  // Only change code below this line

  // Only change code above this line
  return arr;
}

console.log(mixedNumbers(['IV', 5, 'six']));
```

# --solutions--

```js
function mixedNumbers(arr) {
  arr.push(7,'VIII',9);
  arr.unshift('I',2,'three');
  return arr;
}
```
