---
id: 587d7b7d367417b2b2512b1f
title: 객체에 저장된 배열 수정하기
challengeType: 1
forumTopicId: 301163
dashedName: modify-an-array-stored-in-an-object
---

# --description--

이제 JavaScript 객체의 기본 작업을 모두 살펴보았습니다. 키-값 쌍을 추가, 수정, 삭제할 수 있고, 키가 존재하는지 확인하며, 객체 내 모든 키를 순회할 수 있습니다. JavaScript 학습을 계속하면 객체의 훨씬 더 다양한 활용법을 보게 될 것입니다. 또한 커리큘럼의 코딩 인터뷰 준비 섹션에 있는 자료 구조 수업에서는 ES6의 <dfn>Map</dfn>과 <dfn>Set</dfn> 객체도 다루는데, 이들은 일반 객체와 비슷하지만 추가 기능을 제공합니다. 배열과 객체의 기본을 배웠으니 이제 JavaScript를 사용해 더 복잡한 문제를 해결할 준비가 되었습니다!

# --instructions--

코드 편집기에 제공된 객체를 살펴보세요. `user` 객체는 세 개의 키를 포함합니다. `data` 키는 다섯 개의 키를 포함하며, 그중 하나는 `friends` 배열을 포함합니다. 이를 통해 객체가 얼마나 유연한 자료 구조인지 알 수 있습니다. `addFriend` 함수를 작성하기 시작했습니다. `user` 객체를 인수로 받아 `friend` 인수의 이름을 `user.data.friends`에 저장된 배열에 추가하고 그 배열을 반환하도록 완성하세요.

# --hints--

`user` 객체는 `name`, `age`, `data` 키를 가져야 합니다.

```js
assert('name' in user && 'age' in user && 'data' in user);
```

`addFriend` 함수는 `user` 객체와 `friend` 문자열을 인수로 받아 `friends` 객체의 `user` 배열에 친구를 추가해야 합니다.

```js
assert(
  (function () {
    let L1 = user.data.friends.length;
    addFriend(user, 'Sean');
    let L2 = user.data.friends.length;
    return L2 === L1 + 1;
  })()
);
```

`addFriend(user, "Pete")`는 `["Sam", "Kira", "Tomo", "Pete"]`를 반환해야 합니다.

```js
assert.deepEqual(
  (function () {
    delete user.data.friends;
    user.data.friends = ['Sam', 'Kira', 'Tomo'];
    return addFriend(user, 'Pete');
  })(),
  ['Sam', 'Kira', 'Tomo', 'Pete']
);
```

# --seed--

## --seed-contents--

```js
let user = {
  name: 'Kenneth',
  age: 28,
  data: {
    username: 'kennethCodesAllDay',
    joinDate: 'March 26, 2016',
    organization: 'freeCodeCamp',
    friends: [
      'Sam',
      'Kira',
      'Tomo'
    ],
    location: {
      city: 'San Francisco',
      state: 'CA',
      country: 'USA'
    }
  }
};

function addFriend(userObj, friend) {
  // Only change code below this line

  // Only change code above this line
}

console.log(addFriend(user, 'Pete'));
```

# --solutions--

```js
let user = {
  name: 'Kenneth',
  age: 28,
  data: {
    username: 'kennethCodesAllDay',
    joinDate: 'March 26, 2016',
    organization: 'freeCodeCamp',
    friends: [
      'Sam',
      'Kira',
      'Tomo'
    ],
    location: {
      city: 'San Francisco',
      state: 'CA',
      country: 'USA'
    }
  }
};

function addFriend(userObj, friend) {
  userObj.data.friends.push(friend);
  return userObj.data.friends;
}
```
