---
id: 587d78b2367417b2b2512b10
title: splice()를 사용해 항목 삭제하기
challengeType: 1
forumTopicId: 301166
dashedName: remove-items-using-splice
---

# --description--

자, 이제 `shift()`와 `pop()`을 사용해 배열의 처음과 끝에서 요소를 삭제하는 방법을 배웠는데, 만약 중간 어딘가에서 요소를 삭제하거나 한 번에 여러 요소를 삭제하고 싶다면 어떻게 해야 할까요? 바로 그럴 때 `splice()`가 필요합니다. `splice()`는 배열 어디에서든 **연속된 여러 요소를 삭제**할 수 있게 해줍니다.

`splice()`는 최대 3개의 매개변수를 받을 수 있지만, 지금은 첫 2개에만 집중하겠습니다. `splice()`의 첫 두 매개변수는 정수로, `splice()`가 호출되는 배열 내 항목의 인덱스 또는 위치를 나타냅니다. 그리고 배열은 *0부터 인덱스가 시작*하므로, 배열의 첫 번째 요소를 가리키려면 `0`를 사용합니다. `splice()`의 첫 번째 매개변수는 요소 삭제를 시작할 배열의 인덱스를 나타내고, 두 번째 매개변수는 삭제할 요소의 개수를 나타냅니다. 예를 들어:

```js
let array = ['today', 'was', 'not', 'so', 'great'];

array.splice(2, 2);
```

여기서는 세 번째 요소(인덱스 2)부터 시작해 2개의 요소를 삭제합니다. `array`는 `['today', 'was', 'great']` 값을 갖게 됩니다.

`splice()`는 호출된 배열을 수정할 뿐 아니라, 삭제된 요소들의 값을 담은 새 배열도 반환합니다:

```js
let array = ['I', 'am', 'feeling', 'really', 'happy'];

let newArray = array.splice(3, 2);
```

`newArray`는 `['really', 'happy']` 값을 갖습니다.

# --instructions--

배열 `arr`가 초기화되어 있습니다. `splice()`를 사용해 `arr`에서 요소를 삭제하여, 합이 `10`인 요소만 남도록 하세요.

# --hints--

원래 `const arr = [2, 4, 5, 1, 7, 5, 2, 1];` 줄은 변경하지 마세요.

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/constarr=\[2,4,5,1,7,5,2,1\];?/)
);
```

`arr`에는 합이 `10`인 요소만 있어야 합니다.

```js
assert.strictEqual(
  arr.reduce((a, b) => a + b),
  10
);
```

코드는 `splice()`에서 `arr` 메서드를 사용해야 합니다.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/arr\.splice\(/));
```

splice는 `arr`에서만 요소를 삭제하고 `arr`에 추가 요소를 넣지 않아야 합니다.

```js
assert(
  !__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/arr\.splice\(\d+,\d+,\d+.*\)/g)
);
```

# --seed--

## --seed-contents--

```js
const arr = [2, 4, 5, 1, 7, 5, 2, 1];
// Only change code below this line

// Only change code above this line
console.log(arr);
```

# --solutions--

```js
const arr = [2, 4, 5, 1, 7, 5, 2, 1];
arr.splice(1, 4);
```
