---
id: 587d7b7e367417b2b2512b20
title: 배열을 사용해 데이터 모음을 저장하기
challengeType: 1
forumTopicId: 301167
dashedName: use-an-array-to-store-a-collection-of-data
---

# --description--

아래는 가장 간단한 배열 자료 구조 구현 예시입니다. 이것은 <dfn>일차원 배열</dfn>로, 한 단계만 있거나 다른 배열이 중첩되어 있지 않음을 의미합니다. 여기에는 <dfn>불리언</dfn>, <dfn>문자열</dfn>, <dfn>숫자</dfn> 등 유효한 JavaScript 데이터 유형이 포함되어 있음을 확인하세요.

```js
let simpleArray = ['one', 2, 'three', true, false, undefined, null];
console.log(simpleArray.length);
```

`console.log` 호출은 `7`을 표시합니다.

모든 배열은 길이 속성을 가지며, 위에서 보듯이 `Array.length` 구문으로 매우 쉽게 접근할 수 있습니다. 아래는 더 복잡한 배열 구현 예시입니다. 이것은 <dfn>다차원 배열</dfn>로, 다른 배열을 포함하는 배열입니다. 이 배열에는 JavaScript <dfn>객체</dfn>도 포함되어 있는데, 다음 섹션에서 자세히 살펴보겠지만 지금은 배열이 복잡한 객체도 저장할 수 있다는 것만 알아두세요.

```js
let complexArray = [
  [
    {
      one: 1,
      two: 2
    },
    {
      three: 3,
      four: 4
    }
  ],
  [
    {
      a: "a",
      b: "b"
    },
    {
      c: "c",
      d: "d"
    }
  ]
];
```

# --instructions--

`yourArray`라는 변수를 정의했습니다. `yourArray` 변수에 길이가 최소 5인 배열을 할당하여 문장을 완성하세요. 배열에는 최소 하나의 <dfn>문자열</dfn>, 하나의 <dfn>숫자</dfn>, 하나의 <dfn>불리언</dfn>이 포함되어야 합니다.

# --hints--

`yourArray`는 배열이어야 합니다.

```js
assert.strictEqual(Array.isArray(yourArray), true);
```

`yourArray`는 최소 5개의 요소를 가져야 합니다.

```js
assert.isAtLeast(yourArray.length, 5);
```

`yourArray`에는 최소 하나의 `boolean`가 포함되어야 합니다.

```js
assert(yourArray.filter((el) => typeof el === 'boolean').length >= 1);
```

`yourArray`에는 최소 하나의 `number`가 포함되어야 합니다.

```js
assert(yourArray.filter((el) => typeof el === 'number').length >= 1);
```

`yourArray`에는 최소 하나의 `string`가 포함되어야 합니다.

```js
assert(yourArray.filter((el) => typeof el === 'string').length >= 1);
```

# --seed--

## --seed-contents--

```js
let yourArray; // Change this line
```

# --solutions--

```js
let yourArray = ['a string', 100, true, ['one', 2], 'another string'];
```
