---
id: 587d7b7c367417b2b2512b1a
title: 대괄호 표기법으로 속성 이름 접근하기
challengeType: 1
forumTopicId: 301150
dashedName: access-property-names-with-bracket-notation
---

# --description--

첫 번째 객체 도전 과제에서 변수 평가를 사용해 속성 값을 접근하는 방법으로 대괄호 표기법을 언급했습니다. 예를 들어, `foods` 객체가 슈퍼마켓 계산대 프로그램에서 사용된다고 상상해 보세요. `selectedFood`를 설정하는 함수가 있고, `foods` 객체에서 해당 식품이 있는지 확인하려고 합니다. 이 코드는 다음과 같을 수 있습니다:

```js
let selectedFood = getCurrentFood(scannedItem);
let inventory = foods[selectedFood];
```

이 코드는 `selectedFood` 변수에 저장된 값을 평가하고 `foods` 객체에서 그 키의 값을 반환하거나, 없으면 `undefined`를 반환합니다. 대괄호 표기법은 객체 속성이 런타임 전에 알려지지 않았거나 더 동적으로 접근해야 할 때 매우 유용합니다.

# --instructions--

`checkInventory`라는 함수를 정의했습니다. 이 함수는 스캔된 항목을 인수로 받습니다. `scannedItem` 객체에서 `foods` 키의 현재 값을 반환하세요. `checkInventory`에 유효한 키만 인수로 제공된다고 가정해도 됩니다.

# --hints--

`checkInventory`은 함수여야 합니다.

```js
assert.strictEqual(typeof checkInventory, 'function');
```

`foods` 객체는 다음 키-값 쌍만 가져야 합니다: `apples: 25`, `oranges: 32`, `plums: 28`, `bananas: 13`, `grapes: 35`, `strawberries: 27`.

```js
assert.deepEqual(foods, {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
});
```

`checkInventory("apples")`은 `25`를 반환해야 합니다.

```js
assert.strictEqual(checkInventory('apples'), 25);
```

`checkInventory("bananas")`은 `13`를 반환해야 합니다.

```js
assert.strictEqual(checkInventory('bananas'), 13);
```

`checkInventory("strawberries")`는 `27`를 반환해야 합니다.

```js
assert.strictEqual(checkInventory('strawberries'), 27);
```

# --seed--

## --seed-contents--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

function checkInventory(scannedItem) {
  // Only change code below this line

  // Only change code above this line
}

console.log(checkInventory("apples"));
```

# --solutions--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

function checkInventory(scannedItem) {
  return foods[scannedItem];
}
```
