---
id: 587d78b3367417b2b2512b11
title: splice()를 사용해 항목 추가하기
challengeType: 1
forumTopicId: 301152
dashedName: add-items-using-splice
---

# --description--

지난 도전 과제에서 `splice()`가 최대 세 개의 매개변수를 받을 수 있다고 언급한 것을 기억하세요? 세 번째 매개변수로 하나 이상의 요소를 포함해 배열에 추가할 수 있습니다. 이는 요소 하나 또는 여러 개를 다른 것으로 빠르게 교체할 때 매우 유용합니다.

```js
const numbers = [10, 11, 12, 12, 15];
const startIndex = 3;
const amountToDelete = 1;

numbers.splice(startIndex, amountToDelete, 13, 14);
console.log(numbers);
```

`12`가 두 번째로 나타난 부분은 삭제되고, 같은 인덱스에 `13`와 `14`가 추가됩니다. `numbers` 배열은 이제 `[ 10, 11, 12, 13, 14, 15 ]`가 됩니다.

여기서는 숫자 배열로 시작합니다. 그런 다음 `splice()`에 다음을 전달합니다: 삭제를 시작할 인덱스(3), 삭제할 요소 수(1), 그리고 나머지 인수(13, 14)는 같은 인덱스부터 삽입됩니다. `amountToDelete` 뒤에는 쉼표로 구분된 여러 요소가 올 수 있으며, 각각 삽입됩니다.

# --instructions--

HTML 색상 배열을 인수로 받는 `htmlColorNames` 함수를 정의했습니다. `splice()`를 사용해 배열의 처음 두 요소를 삭제하고, 그 자리에 `'DarkSalmon'`과 `'BlanchedAlmond'`을 각각 추가하도록 함수를 수정하세요.

# --hints--

`htmlColorNames`는 `["DarkSalmon", "BlanchedAlmond", "LavenderBlush", "PaleTurquoise", "FireBrick"]`를 반환해야 합니다.

```js
assert.deepEqual(
  htmlColorNames([
    'DarkGoldenRod',
    'WhiteSmoke',
    'LavenderBlush',
    'PaleTurquoise',
    'FireBrick'
  ]),
  [
    'DarkSalmon',
    'BlanchedAlmond',
    'LavenderBlush',
    'PaleTurquoise',
    'FireBrick'
  ]
);
```

`htmlColorNames` 함수는 `splice()` 메서드를 사용해야 합니다.

```js
assert(/.splice/.test(__helpers.removeJSComments(code)));
```

`shift()`나 `unshift()`를 사용하지 마세요.

```js
assert(!/shift|unshift/.test(__helpers.removeJSComments(code)));
```

배열 대괄호 표기법을 사용하지 마세요.

```js
assert(!/\[\d\]\s*=/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function htmlColorNames(arr) {
  // Only change code below this line

  // Only change code above this line
  return arr;
}

console.log(htmlColorNames(['DarkGoldenRod', 'WhiteSmoke', 'LavenderBlush', 'PaleTurquoise', 'FireBrick']));
```

# --solutions--

```js
function htmlColorNames(arr) {
  arr.splice(0,2,'DarkSalmon', 'BlanchedAlmond');
  return arr;
}
```
