---
id: 587d7b7c367417b2b2512b1b
title: 객체 프로퍼티를 삭제하려면 delete 키워드를 사용하세요
challengeType: 1
forumTopicId: 301168
dashedName: use-the-delete-keyword-to-remove-object-properties
---

# --description--

이제 객체가 무엇이고 기본적인 기능과 장점이 무엇인지 알게 되었습니다. 요약하자면, 객체는 데이터를 구조화하는 데 유연하고 직관적인 방법을 제공하는 키-값 저장소이며, 매우 빠른 조회 시간을 제공합니다. 이 도전 과제들 전반에 걸쳐 객체에 대해 수행할 수 있는 여러 일반적인 작업을 설명할 것이며, 이를 통해 프로그램에서 이 유용한 자료 구조를 편하게 사용할 수 있게 될 것입니다.

이전 도전 과제에서는 객체의 키-값 쌍을 추가하거나 수정하는 작업을 했습니다. 여기서는 객체에서 키-값 쌍을 *삭제*하는 방법을 살펴보겠습니다.

마지막으로 `foods` 객체 예제를 다시 살펴보겠습니다. `apples` 키를 삭제하고 싶다면, 다음과 같이 `delete` 키워드를 사용해 삭제할 수 있습니다:

```js
delete foods.apples;
```

# --instructions--

delete 키워드를 사용해 `oranges`, `plums`, `strawberries` 키를 `foods` 객체에서 삭제하세요.

# --hints--

`foods` 객체에는 `apples`, `grapes`, `bananas` 세 개의 키만 있어야 합니다.

```js
assert(
  !foods.hasOwnProperty('oranges') &&
    !foods.hasOwnProperty('plums') &&
    !foods.hasOwnProperty('strawberries') &&
    Object.keys(foods).length === 3
);
```

`oranges`, `plums`, `strawberries` 키는 `delete`를 사용해 삭제해야 합니다.

```js
assert(
  __helpers.removeJSComments(code).search(/oranges:/) !== -1 &&
    __helpers.removeJSComments(code).search(/plums:/) !== -1 &&
    __helpers.removeJSComments(code).search(/strawberries:/) !== -1
);
```

# --seed--

## --seed-contents--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

// Only change code below this line

// Only change code above this line

console.log(foods);
```

# --solutions--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

delete foods.oranges;
delete foods.plums;
delete foods.strawberries;

console.log(foods);
```
