---
id: 587d7b7b367417b2b2512b14
title: indexOf()로 요소 존재 여부 확인하기
challengeType: 1
forumTopicId: 301154
dashedName: check-for-the-presence-of-an-element-with-indexof
---

# --description--

배열은 언제든지 변경되거나 *변이*될 수 있기 때문에, 특정 데이터가 배열 내 어디에 있을지, 또는 그 요소가 여전히 존재하는지 보장할 수 없습니다. 다행히 JavaScript는 배열에 요소가 존재하는지 빠르고 쉽게 확인할 수 있는 또 다른 내장 메서드 `indexOf()`를 제공합니다. `indexOf()`는 요소를 매개변수로 받고, 호출하면 그 요소의 위치, 즉 인덱스를 반환하며, 배열에 요소가 없으면 `-1`를 반환합니다.

예를 들어:

```js
let fruits = ['apples', 'pears', 'oranges', 'peaches', 'pears'];

fruits.indexOf('dates');
fruits.indexOf('oranges');
fruits.indexOf('pears');
```

`indexOf('dates')`는 `-1`를 반환하고, `indexOf('oranges')`는 `2`를 반환하며, `indexOf('pears')`는 `1`를 반환합니다(각 요소가 존재하는 첫 번째 인덱스).

# --instructions--

`indexOf()`는 배열에 요소가 존재하는지 빠르게 확인할 때 매우 유용합니다. 배열과 요소를 인수로 받는 함수 `quickCheck`가 정의되어 있습니다. `indexOf()`를 사용해 함수를 수정하여, 전달된 요소가 배열에 존재하면 `true`를 반환하고, 존재하지 않으면 `false`를 반환하게 하세요.

# --hints--

`quickCheck` 함수는 문자열(`true` 또는 `false`)이 아니라 불리언(`"true"` 또는 `"false"`)을 반환해야 합니다.

```js
assert.isBoolean(quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'));
```

`quickCheck(["squash", "onions", "shallots"], "mushrooms")`는 `false`를 반환해야 합니다.

```js
assert.strictEqual(
  quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'),
  false
);
```

`quickCheck(["onions", "squash", "shallots"], "onions")`는 `true`를 반환해야 합니다.

```js
assert.strictEqual(
  quickCheck(['onions', 'squash', 'shallots'], 'onions'),
  true
);
```

`quickCheck([3, 5, 9, 125, 45, 2], 125)`는 `true`를 반환해야 합니다.

```js
assert.strictEqual(quickCheck([3, 5, 9, 125, 45, 2], 125), true);
```

`quickCheck([true, false, false], undefined)`는 `false`를 반환해야 합니다.

```js
assert.strictEqual(quickCheck([true, false, false], undefined), false);
```

`quickCheck` 함수는 `indexOf()` 메서드를 활용해야 합니다.

```js
assert.notStrictEqual(quickCheck.toString().search(/\.indexOf\(/), -1);
```

# --seed--

## --seed-contents--

```js
function quickCheck(arr, elem) {
  // Only change code below this line

  // Only change code above this line
}

console.log(quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'));
```

# --solutions--

```js
function quickCheck(arr, elem) {
  return arr.indexOf(elem) >= 0; 
}
```
