---
id: 5a661e0f1068aca922b3ef17
title: 대괄호 표기법을 사용해 배열 내용에 접근하기
challengeType: 1
forumTopicId: 301149
dashedName: access-an-arrays-contents-using-bracket-notation
---

# --description--

어떤 자료 구조든 기본 기능은 데이터를 저장할 뿐만 아니라 명령에 따라 그 데이터를 꺼내 쓸 수 있어야 한다는 점입니다. 이제 배열을 만드는 방법을 배웠으니, 배열 안의 정보를 어떻게 접근할 수 있을지 생각해 봅시다.

아래처럼 간단한 배열을 정의하면, 3개의 항목이 들어 있습니다:

```js
let ourArray = ["a", "b", "c"];
```

배열에서는 각 항목마다 <dfn>인덱스</dfn>가 있습니다. 이 인덱스는 배열 내 항목의 위치이자 참조 방법입니다. 하지만 JavaScript 배열은 <dfn>0부터 시작하는 인덱스</dfn>를 사용한다는 점이 중요합니다. 즉, 배열의 첫 번째 요소는 실제로 첫 번째 위치가 아니라 ***0번째*** 위치에 있습니다. 배열에서 요소를 꺼내려면 인덱스를 괄호로 감싸 배열 끝에 붙이거나, 보통은 배열 객체를 참조하는 변수 뒤에 붙입니다. 이것을 <dfn>대괄호 표기법</dfn>이라고 합니다. 예를 들어, `a`에서 `ourArray`를 꺼내 변수에 할당하려면 다음 코드를 사용합니다:

```js
let ourVariable = ourArray[0];
```

이제 `ourVariable`는 `a` 값을 갖게 됩니다.

인덱스에 연결된 값을 접근하는 것 외에도, 같은 표기법으로 인덱스에 값을 *할당*할 수도 있습니다:

```js
ourArray[1] = "not b anymore";
```

대괄호 표기법을 사용해 인덱스 1에 있던 항목을 문자열 `b`에서 `not b anymore`로 바꿨습니다. 이제 `ourArray`는 `["a", "not b anymore", "c"]`가 되었습니다.

# --instructions--

이 도전 과제를 완료하려면, `1`의 두 번째 위치(인덱스 `myArray`)에 `b` 문자를 제외한 원하는 값을 할당하세요.

# --hints--

`myArray[0]`는 문자 `a`와 같아야 합니다.

```js
assert.strictEqual(myArray[0], 'a');
```

`myArray[1]`는 문자 `b`와 같지 않아야 합니다.

```js
assert.notStrictEqual(myArray[1], 'b');
```

`myArray[2]`는 문자 `c`와 같아야 합니다.

```js
assert.strictEqual(myArray[2], 'c');
```

`myArray[3]`는 문자 `d`와 같아야 합니다.

```js
assert.strictEqual(myArray[3], 'd');
```

# --seed--

## --seed-contents--

```js
let myArray = ["a", "b", "c", "d"];
// Only change code below this line

// Only change code above this line
console.log(myArray);
```

# --solutions--

```js
let myArray = ["a", "b", "c", "d"];
myArray[1] = "e";
```
