---
id: 680a836235461fce1ff53b3f
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

`translate()` メソッドは、`str.maketrans()` によって生成された変換テーブルを引数として受け取ります。このメソッドは文字列に対して呼び出され、変換テーブルに基づいて文字が置き換えられた元の文字列のコピーを返します。

```py
t = str.maketrans('lk', 'br')
sentence = 'The tent gave in to the leaks.'

print(sentence.translate(t))
# Output: The tent gave in to the bears.
```

`translate()` に対して `text` メソッドを呼び出し、`translation_table` を引数として渡し、その結果を `encrypted_text` という名前の変数に代入してください。

# --hints--

`encrypted_text` という名前の変数が必要です。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

`text.translate(translation_table)` を `encrypted_text` に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
--fcc-editable-region--

--fcc-editable-region--
```
