---
id: 680a5e7f3d47715196b12ca0
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

シフトしたアルファベット全体ができたので、`print(shifted_alphabet)` の呼び出しを削除してください。

# --hints--

コードから `print(shifted_alphabet)` を削除してください。

```js
({ test: () => assert(runPython(`not _Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
print(shifted_alphabet)
--fcc-editable-region--
```
