---
id: 680a7e200400089967b8b1f4
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

`str.maketrans()` メソッドは、同じ長さの2つの文字列を受け取り、最初の文字列の各文字を2番目の文字列の対応する文字にマッピングする翻訳テーブルを返します。翻訳テーブル内の各文字は Unicode 順位として保存されており、これは文字を一意に識別する番号です。

```py
lower_chars = 'abc'
upper_chars = 'ABC'

table = str.maketrans(lower_chars, upper_chars)
print(table) # {97: 65, 98: 66, 99: 67}
```

カリキュラムの後半で`str.maketrans()`が返す構造の種類について詳しく学びます。今は、`alphabet`の文字を`shifted_alphabet`の文字にマッピングする変換テーブルを作成し、`translation_table`という名前の変数に代入してください。

# --hints--

`translation_table`という名前の変数を宣言してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("translation_table")`)) })
```

`str.maketrans()`を呼び出し、`alphabet`と`shifted_alphabet`を引数として渡し、その結果を`translation_table`に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("translation_table").is_equivalent("translation_table = str.maketrans(alphabet, shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--

--fcc-editable-region--
```
