---
id: 680a549f60f5a3fedfc0edd9
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

出力からわかるように、シフトされたアルファベットは `f` から始まります。なぜなら `shift` の値が `5` だからです。しかし、シフトされたアルファベットからは最初の5文字、`a`、`b`、`c`、`d`、`e` が抜けているので、これらをシフトされたアルファベットの末尾に追加する必要があります。

`+` 演算子は、連結と呼ばれる処理で2つ以上の文字列を結合するために使われます。例えば次のように使います。

```py
greeting = 'Hello' + ' ' + 'World'
print(greeting) # Hello World

```

既存の `shifted_alphabet` 変数の代入を修正してください。スライス構文を使って `alphabet` の欠けている最初の部分を抽出し、それを `alphabet[shift:]` に連結してください。

念のために説明すると、`sentence[start:stop]` は `sentence` の位置 `start`（含む）から `stop`（含まない）までの文字を返します。

# --hints--

`alphabet` 変数を直接変更しないでください。

```js
({ test: () => runPython(`
  assert alphabet == "abcdefghijklmnopqrstuvwxyz"
`) })
```

`alphabet[shift:]` と `alphabet` の欠けている最初の部分の連結を、宣言された同じ行で `shifted_alphabet` 変数に代入してください。スライスの終了位置を指定するには `shift` を使ってください。

```js
({ test: () => runPython(`
_shifted_alpha = _Node(_code).find_variable("shifted_alphabet")
_first = _shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[:shift]")
_second =_shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[0:shift]")
assert _first or _second
`) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--
shifted_alphabet = alphabet[shift:]
--fcc-editable-region--
print(shifted_alphabet)
```
