---
id: 680a82806c4c6bc3567f2708
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

`text` という名前の新しい変数を宣言し、文字列 `hello world` を代入してください。これは暗号化するメッセージになります。

# --hints--

`text` という名前の変数を宣言してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("text")`)) })
```

文字列 `hello world` を `text` に代入してください。文字列はシングルクォートか二重引用符で囲み、文字の大文字・小文字に注意してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("text").is_equivalent("text = 'hello world'")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--

--fcc-editable-region--
```
