---
id: 680a41ef5bfdc059cfd9d4b5
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

前のレッスンでは、変数にストアできるさまざまなデータ型について学びました。整数値を代入したばかりです。今度は文字列を代入する必要があります。文字列とは、シングルクォートか二重引用符で囲まれた文字のシーケンスです。

```py
string_1 = 'I am a string'
string_2 = "I am also a string"
```

`alphabet` という名前の別の変数を宣言し、この変数に `abcdefghijklmnopqrstuvwxyz` という文字列を代入してください。

# --hints--

`alphabet` という名前の変数を宣言してください。変数名は行の先頭に置くよう注意してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("alphabet")`)) })
```

`abcdefghijklmnopqrstuvwxyz` 変数に `alphabet` という文字列を代入してください。文字列を囲むにはシングルクォートか二重引用符のどちらかを使うことを忘れないでください。

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable("alphabet").is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
shift = 5
```
