---
id: 680a3d4c065557260de2ff33
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このワークショップでは、アルファベットの文字をずらしてメッセージを暗号化する方法であるシーザー暗号というプログラムを作成します。

まず、`shift` という名前の変数を作成し、新しい変数に `5` の値を代入してください。

# --hints--

`shift` という名前の変数を宣言してください。変数名は行の先頭に置くよう注意しましょう。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

`5` の値を `shift` 変数に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
