---
id: 680a431692a5f863f726c438
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

このワークショップでは、*シーザー暗号*を作成します。これはテキストを暗号化する最も簡単な手法の一つで、平文の各文字をアルファベット上で一定の位置だけずらした文字に置き換える方法です。例えば、シフトが5の場合、`a` は `f` に、`b` は `g` に置き換えられます。

この暗号を実装するには、シフトで示された位置から始まる新しいアルファベットのバージョンを作成する必要があります。前のレッスンで学んだように、文字列の一部を文字列スライスで取り出せます。

```py
fcc = 'freeCodeCamp'
print(fcc[8:]) # Camp
```

`shifted_alphabet` という名前の変数を作成し、スライス構文を使って `alphabet` の `shift` のインデックスから始まる部分を代入してください。その後、組み込みの `print()` 関数を呼び出して `shifted_alphabet` をターミナルに出力し、結果を確認しましょう。

# --hints--

`shifted_alphabet` という名前の新しい変数を宣言してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shifted_alphabet")`)) })

```

`alphabet[shift:]` に `shifted_alphabet` を代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shifted_alphabet").is_equivalent("shifted_alphabet = alphabet[shift:]")`)) })
```

`print()` 関数を呼び出し、`shifted_alphabet` を引数として渡してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--

--fcc-editable-region--
```
