---
id: 6810891bd6de8a87a833df42
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

出力からわかるように、メッセージは暗号化されています。次のステップでは、異なるメッセージを暗号化したい場合に備えて、コードを再利用できるようにします。

そのためには、関数を作成する必要があります。復習として、`spam` という名前の関数を作成し、ターミナルに `Spam!` を出力する方法は以下の通りです。

```py
def spam():
    print('Spam!')
```

`caesar` という名前の関数を作成してください。既存のコードをすべて関数の本体内に入れます。関数本体内のすべての行で同じインデントレベルを保つように注意してください。

# --hints--

`caesar` という名前の関数があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).has_function("caesar")`)) })
```

これまで書いたコードをすべて `caesar` 関数の本体内に移動してください。関数本体内のすべての行で同じインデントレベルを保つようにしてください...

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_body().is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'\\nshift = 5\\nshifted_alphabet = alphabet[shift:] + alphabet[:shift]\\ntranslation_table = str.maketrans(alphabet, shifted_alphabet)\\ntext = 'hello world'\\nencrypted_text = text.translate(translation_table)\\nprint(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
encrypted_text = text.translate(translation_table)
print(encrypted_text)
--fcc-editable-region--
```
