---
id: 6818b80519e6c6c47a8bb6e8
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

前のレッスンで学んだように、関数はパラメータを持つことができ、これは関数内で参照できる変数です。ここに `sum` 関数があり、2つのパラメータ `num1` と `num2` があります。

```py
def sum(num1, num2):
    print(num1 + num2)
```

暗号化するメッセージとシフトはまだ関数内でハードコードされています。関数に `text` と `shift` という名前の2つのパラメータを与え、`text` 関数の本体から `shift` と `caesar` の両方の変数宣言を削除してください。

# --hints--

`caesar` 関数は `text` と `shift` の2つのパラメータを持つべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift")`)) })
```

`text` 関数の本体内に `caesar` 変数が宣言されていてはいけません。

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("text")`)) })
```

`shift` 関数の本体内に `caesar` 変数が宣言されていてはいけません。

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("shift")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shift = 5
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    text = 'hello world'
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
```
