---
id: 6818eb3f98c0bca5af7c8d03
title: ステップ 13
challengeType: 20
dashedName: step-13
---

# --description--

これでコードが再利用可能になりました。ただし、`caesar` 関数はターミナルにメッセージを出力し、デフォルトでは `None` を返します。それを証明するために、コードの最後に `encrypted_text` を出力してください。

# --hints--

関数本体の外側で `encrypted_text` を出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

encrypted_text = caesar('freeCodeCamp', 3)
--fcc-editable-region--

--fcc-editable-region--
```
