---
id: 6818e9c7c8dc5694e8f64fc3
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

今、`caesar` 関数を、文字列 `freeCodeCamp` と数値 `3` を引数として呼び出してテストしてください。関数の呼び出し結果を `encrypted_text` という名前の変数に代入してください。

# --hints--

グローバルスコープに `encrypted_text` という名前の変数があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

`caesar('freeCodeCamp', 3)` 変数に `encrypted_text` を代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = caesar('freeCodeCamp', 3)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
