---
id: 6818fcb250f34e62cab32c39
title: ステップ 15
challengeType: 20
dashedName: step-15
---

# --description--

メッセージは暗号化されているのに、大文字の文字はそのまま残っていることに気づいたかもしれません。これは、変換テーブルに大文字が含まれていないために起こります。

これを修正するには、`str.maketrans()` の呼び出しに渡す文字列を `upper()` メソッドを使って変更する必要があります。

```py
greet = "Hello"
print(greet.upper()) # HELLO
```

`str.maketrans()` の呼び出しを更新して、各引数にその引数の大文字バージョンを連結してください。

# --hints--

`str.maketrans()` の呼び出しの引数を、それぞれの引数に大文字バージョンを連結するように変更してください。

```js
({ test: () => assert(runPython(`caesar('This Is A TEST', 5) == 'Ymnx Nx F YJXY'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
