---
id: 6818ed1adc249bb91d373729
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

前のレッスンでは、`return` 文を使って関数から値を返し、その値をコードの他の場所で使えるようにすることを学びました。

```py
def spam():
    return 'Spam!'
```

関数から`print(encrypted_text)` の呼び出しを削除してください。それから、`encrypted_text` 変数の宣言を削除し、関数から直接`text.translate(translation_table)` を返してください。

# --hints--

`print(encrypted_text)` 関数内に`caesar` を含めてはいけません。

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_call("print(encrypted_text)")`)) })
```

`encrypted_text` 関数内に`caesar` 変数を含めてはいけません。

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_variable("encrypted_text")`)) })
```

`caesar` 関数は`text.translate(translation_table)` を返すべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_return("text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
