---
id: 6819bc0637b80256a33adccc
title: ステップ 16
challengeType: 20
dashedName: step-16
---

# --description--

暗号の基本的な機能を実装したので、次はバリデーションを追加する時間です。そのために、`if` 文が必要になります。`if` 文の構文をもう一度確認しましょう：

```py
if condition:
    # code to run when condition is true
```

関数本体の最初に、`if` 文を作成してください。今は条件に `True` を使い、`if` 文の中で文字列 `Shift must be an integer value.` を返してください。

# --hints--

`if` 文は `caesar` 関数の中にあるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0]`)) })
```

`if` 文は条件に `True` を使うべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("True")`)) })
```

`Shift must be an integer value.` 文から文字列 `if` を返すべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_bodies()[0].has_return("'Shift must be an integer value.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    
--fcc-editable-region--
    
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
