---
id: 6819ca464db219ab86f37dc4
title: ステップ 20
challengeType: 20
dashedName: step-20
---

# --description--

すでに確認したように、テキストを暗号化するために渡されるシフトは正の整数である必要があります。ただし、`25`の最後のインデックスである`alphabet`を超えてはいけません。

`if`文に、`shift`が`25`より大きいかどうかをチェックする2つ目の条件を追加してください。Pythonでの論理和演算は`or`演算子で実装されていることを覚えておいてください。

また、返されるメッセージを`Shift must be an integer between 1 and 25.`に更新してください。

# --hints--

`or`演算子を使って、既存の条件を変更せずに`shift < 1`の後に、`shift`が`25`より大きいかどうかをチェックする2つ目の条件を追加してください。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift > 25", "25 < shift", "shift >= 26", "26 =< shift"]
assert any(_cond.is_equivalent(f"shift < 1 or {option}") for option in _options)
`) })
```

2つ目の`if`文は`Shift must be an integer between 1 and 25.`を返すべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be an integer between 1 and 25.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'
    
--fcc-editable-region--
    if shift < 1:
        return 'Shift must be a positive integer.'
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
