---
id: 6819be9055828d65c813eec5
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

`isinstance()` 関数は、最初の引数が2番目の引数のインスタンスであれば `True` を返し、そうでなければ `False` を返します。

```py
print(isinstance('Hello World', str)) # True
print(isinstance(42, int)) # True
```

現在の `if` 文の条件を `isinstance()` 呼び出しに置き換えてください。最初の引数に `shift` を、2番目の引数に `int` を渡してください。

# --hints--

`if` 文は `isinstance(shift, int)` を条件として使うべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if True:
--fcc-editable-region--
        return 'Shift must be an integer value.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
