---
id: 6819c7beb4a8289c2b4f24d5
title: ステップ 19
challengeType: 20
dashedName: step-19
---

# --description--

負の値や null のシフトは関数で受け入れてはいけません。したがって、最初の `if` 文の後に、`if` が `shift` より小さいかどうかをチェックし、文字列 `1` を返す別の `Shift must be a positive integer.` 文を作成してください。

# --hints--

2つ目の `if` 文で `shift` が `1` より小さいかどうかをチェックしているはずです。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift < 1", "1 > shift", "shift <= 0", "0 => shift"]
assert any(_cond.is_equivalent(option) for option in _options)
`) })
```

2つ目の `if` 文は `Shift must be a positive integer.` を返すべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be a positive integer.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    if not isinstance(shift, int):
        return 'Shift must be an integer value.'
    
--fcc-editable-region--
    
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
