---
id: 6819c5c9faf2548e556163d5
title: ステップ 18
challengeType: 20
dashedName: step-18
---

# --description--

前のレッスンで学んだように、`not` 演算子は単項論理演算子で、式を否定します。

```py
print(not True) # False
print(not False) # True
```

`not` 演算子を使って、`if` 文の条件を修正し、`Shift must be an integer value.` が整数で *ない* 場合に関数が `shift` を返すようにしてください。

# --hints--

`not` 文の中で `isinstance(shift, int)` を否定するために `if` 演算子を使うべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if isinstance(shift, int):
--fcc-editable-region--
        return 'Shift must be an integer value.'
    
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
