---
id: 6819cf38880d8cc9b9f6af7a
title: ステップ 21
challengeType: 20
dashedName: step-21
---

# --description--

Pythonでは、関数のパラメータにデフォルト値を指定できるため、引数を少なくしたり、まったく渡さずに呼び出せる関数を作成できます。デフォルト値を持つ`name`パラメータを持つ関数を作成する方法は次のとおりです。

```py
def greet(name='Polly'):
    return 'Hello ' + name
    
print(greet()) # Hello Polly
```

関数に`encrypt`という名前の3番目のパラメータを追加し、デフォルト値を`True`にしてください。

# --hints--

関数に`encrypt`という名前の3番目のパラメータを、デフォルト値`True`で追加してください。既存のパラメータは変更しないでください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift, encrypt=True")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
--fcc-editable-region--

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
