---
id: 681a44f4c3235f7d8f428545
title: ステップ 25
challengeType: 20
dashedName: step-25
---

# --description--

これから`decrypt`関数をテストします。`encrypted_text`に代入されている値を、復号すべきメッセージを表す次の文字列に置き換えてください：`Pbhentr vf sbhaq va hayvxryl cynprf.`

次に、`decrypted_text`という名前の変数を宣言し、`decrypt`を最初の引数、`encrypted_text`のシフトを2番目の引数として`13`を呼び出した結果を代入してください。

最後に、`decrypted_text`をターミナルに出力してください。これでシーザー暗号は完成です。

# --hints--

`Pbhentr vf sbhaq va hayvxryl cynprf.`を`encrypted_text`に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = 'Pbhentr vf sbhaq va hayvxryl cynprf.'")`)) })
```

`decrypted_text`という名前の変数を用意してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("decrypted_text")`)) })
```

`decrypt(encrypted_text, 13)`を`decrypted_text`に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("decrypted_text").is_equivalent("decrypted_text = decrypt(encrypted_text, 13)")`)) })
```

`decrypted_text`を出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(decrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)
    
def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

--fcc-editable-region--
encrypted_text = encrypt('freeCodeCamp', 3)
print(encrypted_text)
--fcc-editable-region--
```

# --solutions--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)
    
def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

encrypted_text = 'Pbhentr vf sbhaq va hayvxryl cynprf.'
decrypted_text = decrypt(encrypted_text, 13)
print(decrypted_text)
```
