---
id: 681a1d60e037127e7c8e0230
title: ステップ 22
challengeType: 20
dashedName: step-22
---

# --description--

前のステップで追加したパラメータを使って、関数が渡されたテキストを暗号化するか（デフォルトの動作、`encrypt=True`）、または暗号化されたメッセージを復号するかを判断します。

`if`が真でないかどうかをチェックする`encrypt`文を作成してください。新しい`if`文の中で、`shift`を`- shift`に設定します。これは暗号化プロセスに対して逆方向にシフトを行うために必要です。

# --hints--

`shift`が真でない場合、`- shift`を`encrypt`に設定してください。

```js
({ test: () => assert(runPython(`caesar('Ymnx Nx F Yjxy', 5, False) == 'This Is A Test'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    
--fcc-editable-region--
    
--fcc-editable-region--
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
