---
id: 681a2a6cf95e82ca30866364
title: ステップ 23
challengeType: 20
dashedName: step-23
---

# --description--

`encrypt` と `decrypt` という名前の関数を2つ宣言し、それぞれに `text` と `shift` のパラメータを持たせてください。

`encrypt` は暗号化処理に使い、`decrypt` は復号化に使います。これら2つの処理にはわかりやすい名前を付けてください。

両方の新しい関数から `caesar` と `text` を渡して `shift` の呼び出しを返しますが、`False` 関数には3番目の引数として `decrypt` も渡すようにしてください。

# --hints--

`encrypt` という名前の関数を、`text` と `shift` の2つのパラメータで宣言しているはずです。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_args("text, shift")`)) })
```

`encrypt` 関数は `caesar(text, shift)` を返す必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_return("caesar(text, shift)")`)) })
```

`decrypt` という名前の関数を、`text` と `shift` の2つのパラメータで宣言しているはずです。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("decrypt").has_args("text, shift")`)) })
```

`decrypt` 関数は `caesar(text, shift, False)` を返す必要があります。

```js
({ test: () => runPython(`
_n = _Node(_code).find_function("decrypt")
assert _n.has_return("caesar(text, shift, encrypt=False)") or _n.has_return("caesar(text, shift, False)")
`) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

--fcc-editable-region--

--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
