---
id: 68f890b73bbfc9efa7379e7b
title: Tafuta makosa ya kulazimisha aina katika programu yenye hitilafu
challengeType: 26
dashedName: debug-type-coercion-errors-in-a-buggy-app
---

# --description--

Umejiunga tu na duka la ukuzaji wa mtandao la eneo lako, na zoezi lako la kwanza ni kusafisha msimbo wenye hitilafu uliobaki na wasanidi programu waliopita.

Walikuwa wakijaribu kufanya operesheni za kihesabu rahisi, lakini kitu kilienda vibaya, matokeo hayana maana. Zoezi lako ni kukagua, kutafuta makosa, na kurekebisha msimbo ili ufanye kazi kama inavyotarajiwa na uwe rahisi kusoma.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kigezo kinachoitwa `firstResult` kinachoongeza nambari `5` na `10` kwa usahihi ili kutoa thamani `15`.
2. Unapaswa kuwa na kigezo kinachoitwa `secondResult` kinachopunguza `5` kutoka `8` kwa usahihi ili kutoa thamani `3`.
3. Unapaswa kuwa na kigezo kinachoitwa `thirdResult` kinachotoa thamani `6` kwa kuongeza nambari mbili. Badilisha thamani ya kweli au siyo kweli inayotumika sasa katika usemi wa kihesabu na nambari.
4. Unapaswa kuwa na kigezo kinachoitwa `fourthResult` kinachotoa thamani `8` kwa kuongeza nambari mbili. Badilisha thamani ya kweli au siyo kweli inayotumika sasa katika usemi wa kihesabu na nambari.
5. Unapaswa kuwa na kigezo kinachoitwa `fifthResult` kinachozidisha nambari mbili kwa usahihi kupata bidhaa `20`. Badilisha mfuatano wa herufi unaotumika sasa katika usemi wa kihesabu na nambari.
6. Unapaswa kuwa na kigezo kinachoitwa `sixthResult` kinachoongeza nambari mbili kwa usahihi kutoa thamani `22`. Badilisha `null` inayotumika sasa katika usemi wa kihesabu na nambari.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `firstResult`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.firstResult);
```

`firstResult` inapaswa kuwa na thamani ya nambari `15`.

```js
assert.strictEqual(firstResult, 15);
```

Unapaswa kuwa na kigezo kinachoitwa `secondResult`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.secondResult);
```

`secondResult` inapaswa kuwa na thamani ya nambari `3`.

```js
assert.strictEqual(secondResult, 3);
```

Unapaswa kuwa na kigezo kinachoitwa `thirdResult`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.thirdResult);
```

`thirdResult` haipaswi kujumuisha thamani ya kweli au siyo kweli (`true` au `false`).

```js
const explorer = await __helpers.Explorer(code);
const { thirdResult } = explorer.variables;
assert.notInclude(thirdResult?.value?.toString(), "true")
assert.notInclude(thirdResult?.value?.toString(), "false")
```

`thirdResult` inapaswa kuwa na thamani ya nambari `6`.

```js
assert.strictEqual(thirdResult, 6);
```

Unapaswa kuwa na kigezo kinachoitwa `fourthResult`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.fourthResult);
```

`fourthResult` haipaswi kujumuisha thamani ya kweli au siyo kweli (`true` au `false`).

```js
const explorer = await __helpers.Explorer(code);
const { fourthResult } = explorer.variables;
assert.notInclude(fourthResult?.value?.toString(), "true")
assert.notInclude(fourthResult?.value?.toString(), "false")
```

`fourthResult` inapaswa kuwa na thamani ya nambari `8`.

```js
assert.strictEqual(fourthResult, 8);
```

Unapaswa kuwa na kigezo kinachoitwa `fifthResult`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.fifthResult);
```

`fifthResult` haipaswi kujumuisha kuzidisha mfuatano wa herufi (mfano "10" * 2).

```js
const explorer = await __helpers.Explorer(code);
const { fifthResult } = explorer.variables;
const expression = fifthResult?.value?.toString() || '';
assert.notMatch(expression, /(["'`]).*\1\s*\*/);
```

`fifthResult` inapaswa kuwa na thamani ya nambari `20`.

```js
assert.strictEqual(fifthResult, 20);
```

Unapaswa kuwa na kigezo kinachoitwa `sixthResult`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.sixthResult);
```

`sixthResult` haipaswi kujumuisha `null` katika usemi wake wa kihesabu.

```js
const explorer = await __helpers.Explorer(code);
const { sixthResult } = explorer.variables;
assert.notInclude(sixthResult?.value?.toString(), "null");
```

`sixthResult` inapaswa kuwa na thamani ya nambari `22`.

```js
assert.strictEqual(sixthResult, 22);
```

# --seed--

## --seed-contents--

```js
const firstResult = 5 + "10";
console.log(`5 + 10 = ${firstResult}`);

const secondResult = "Eight" - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = true + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = false + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = "10" * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = null + 22;
console.log(`0 + 22 = ${sixthResult}`);
```

# --solutions--

```js
const firstResult = 5 + 10;
console.log(`5 + 10 = ${firstResult}`);

const secondResult = 8 - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = 1 + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = 0 + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = 10 * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = 0 + 22;
console.log(`0 + 22 = ${sixthResult}`);
```
