---
id: 69860157f7fce39f9df399a3
title: 实现 N-Queens 算法
challengeType: 26
dashedName: implement-the-n-queens-algorithm
---

# --description--

N 皇后问题要求你在一个 N×N 的棋盘上放置 N 个皇后，使得没有两个皇后互相攻击（没有两个皇后共享同一行、列或对角线）。 

例如，如果有一个 4x4 的棋盘，一个有效的排列是：

```md
[1, 3, 0, 2]
```

这意味着在第 0 行，皇后放置在第 1 列；在第 1 行，皇后放置在第 3 列；在第 2 行，皇后放置在第 0 列；在第 3 行，皇后放置在第 2 列。

从视觉上看，这种排列方式如下：

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

其中 `Q` 表示一个皇后，`.` 表示一个空白格子。

在本实验中，你将使用深度优先搜索方法实现 N 皇后问题求解器。


**目标**：完成以下用户故事并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个名为 `dfsNQueens` 的函数。
2. 该函数应接受且仅接受一个参数：整数 `n`。
3. 如果 `n` 小于 `1`，该函数应返回一个空数组（`[]`）。
4. 该函数应返回一个方案列表；每个方案本身是一个长度为 `n` 的列表，其中索引为 `i` 的元素是第 `i` 行皇后的列索引（从 0 开始）。

# --hints--

你应该有一个名为 `dfsNQueens` 的函数，它接受一个参数。

```js
assert.isFunction(dfsNQueens);
assert.lengthOf(dfsNQueens, 1);
```

如果 `n` 小于 `1`，该函数应返回一个空数组。

```js
assert.sameDeepOrderedMembers(dfsNQueens(0), []);
assert.sameDeepOrderedMembers(dfsNQueens(-1), []);
assert.sameDeepOrderedMembers(dfsNQueens(-5), []);
```

该函数应返回一个包含多个方案的数组，其中每个方案都是长度为 `n` 的数组。

```js
const output = dfsNQueens(4);
assert.isArray(output);
output.forEach(solution => {
    assert.isArray(solution);
    assert.lengthOf(solution, 4);
})
```

`dfsNQueens(1)` 应该返回 `[[0]]`。

```js
assert.sameDeepOrderedMembers(dfsNQueens(1), [[0]]);
```

`dfsNQueens(2)` 应该返回 `[]`。

```js
assert.sameDeepOrderedMembers(dfsNQueens(2), []);
```

`dfsNQueens(3)` 应该返回 `[]`。

```js
assert.sameDeepOrderedMembers(dfsNQueens(3), []);
```

`dfsNQueens(4)` 应该返回 `[[1, 3, 0, 2], [2, 0, 3, 1]]`。

```js
assert.sameDeepOrderedMembers(dfsNQueens(4), [[1, 3, 0, 2], [2, 0, 3, 1]]);
```

`dfsNQueens(5)` 应该 `返回` `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`。

```js
assert.sameDeepOrderedMembers(dfsNQueens(5), [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]);
```

`dfsNQueens(5).length` 应该是 `10`。

```js
assert.lengthOf(dfsNQueens(5), 10);
```

`dfsNQueens(8).length` 应该是 `92`。

```js
assert.lengthOf(dfsNQueens(8), 92);
```

`dfsNQueens` 应该返回正确的结果。

```js
assert.lengthOf(dfsNQueens(7), 40);
assert.sameDeepOrderedMembers(dfsNQueens(6), [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfsNQueens(n) {
  if (n < 1) return [];

  const results = [];

  function isSafe(queens, row, col) {
    for (let r = 0; r < queens.length; r++) {
      const c = queens[r];
      if (c === col || Math.abs(row - r) === Math.abs(col - c)) {
        return false;
      }
    }
    return true;
  }

  function dfs(queens) {
    const row = queens.length;
    if (row === n) {
      results.push(queens.slice());
      return;
    }

    for (let col = 0; col < n; col++) {
      if (isSafe(queens, row, col)) {
        queens.push(col);
        dfs(queens);
        queens.pop();
      }
    }
  }

  dfs([]);
  return results;
}
```
