---
id: 68c362b379059c388d3874f2
title: 디지털 펫 게임 만들기
challengeType: 25
dashedName: lab-digital-pet-game
demoType: onClick
---

# --description--

**목표:** 아래 사용자 스토리를 완성하고 모든 테스트를 통과하여 도전 과제를 완료하세요.

**사용자 스토리:**

애플리케이션은 두 개의 뷰를 가져야 합니다. 각 뷰에 대한 요구 사항은 나머지 사용자 스토리에 포함되어 있습니다:

첫 번째 뷰:

- 하나의 문자열 입력을 받는 폼을 포함해야 합니다.
- 사용자에게 펫의 이름을 묻는 입력 필드를 제공해야 합니다.
- 게임을 시작하는 제출 버튼을 포함해야 합니다.
- 제출 시 사라져야 합니다.

두 번째 뷰:

- 첫 번째 뷰 제출 후 나타나야 합니다.
- 세 개의 `<button>` 요소를 포함해야 합니다.
- 각 버튼 요소는 다음 세 가지 행동 중 하나와 관련되어야 합니다: `[EAT, PLAY, SLEEP]`.

  - `EAT` 버튼은 펫의 `Hunger` 값을 줄이고, `Energy` 값을 늘려야 합니다.
  - `PLAY` 버튼은 펫의 `Energy` 값을 줄이고, `Happiness` 값을 늘려야 합니다.
  - `SLEEP` 버튼은 펫의 `Hunger` 값을 늘리고, `Energy` 값을 늘려야 합니다.

- 펫의 세 가지 속성을 시각적으로 추적해야 합니다. 이들은 `['Hunger', 'Happiness', 'Energy']`입니다.
- 각 속성을 추적하는 요소는 다음을 만족해야 합니다:

  - 해당 속성과 관련된 텍스트를 포함해야 하며, 대소문자를 구분하지 않습니다. 예: `hunger`, `Hunger`, `hUnGeR`는 모두 유효합니다.
  - 다른 속성과 관련된 텍스트는 포함하지 않아야 합니다. 즉, `Hunger` 추적 요소는 `Happiness` 텍스트를 포함하면 안 됩니다.
  - 해당 속성의 현재 값을 명확한 숫자로 나타내야 합니다.

- 펫이 가진 각 `Action`은 열거자로 감싸져야 합니다.
- 각 `Action` `button`는 해당 버튼 행동과 관련된 다음 `id` 중 하나씩을 포함해야 합니다: `['eat-action', 'play-action', 'sleep-action']`.
- 펫은 `PetMood` 열거자로 감싸진 다음 기분을 가져야 합니다: `['Happy', 'Excited', 'Content', 'Sad', 'Tired', 'Sick', 'Hungry']`.

  - `Hungry`일 때 펫은 `Hunger > 70`여야 합니다.
  - `Tired`일 때 펫은 `Energy < 30`여야 합니다.
  - `Sad`일 때 펫은 `Happiness < 30`여야 합니다.
  - `Excited`이고 `Happiness > 80`일 때 펫은 `Energy > 70`여야 합니다.
  - `Happy`일 때 펫은 `Happiness > 60`여야 합니다.
  - 그렇지 않으면 펫은 `Content`여야 합니다.

- 펫은 각 `Mood`에 연결된 여러 이미지 또는 이모지를 가져야 합니다.

  - 이 이모지 또는 이미지는 현재 `PetMood`에 연결되어 `Record`를 가집니다.
  - 각 `Record` 항목의 키는 관련 `PetMood` 열거형 값이어야 합니다.

# --before-all--

```js
globalThis._PET_NAME = 'Fitzgerald';
globalThis._clock = __FakeTimers.install();
```

# --after-all--

```js
globalThis._clock.uninstall();
delete globalThis._clock;
delete globalThis._PET_NAME;
```

# --hints--

페이지가 로드될 때, `form` 요소가 보여야 합니다.

```js
assert.isNotNull(document.querySelector('form'), 'The form element does not exist at load.');
```

이 폼은 `pet-name` 아이디를 가진 입력 필드를 포함해야 합니다.

```js
assert.isNotNull(document.querySelector('#pet-name'));
```

이 폼은 `button`를 포함해야 합니다.

```js
assert.isNotNull(document.querySelector('button'));
```

입력 필드에 값이 입력되고 버튼이 눌리면 폼은 더 이상 보이지 않아야 합니다.

```js
document.querySelector('#pet-name').value = _PET_NAME;
document.querySelector('button').click();
await _clock.tickAsync(500);

assert.isNull(document.querySelector('form'));
```

두 번째 뷰는 `pet-name` 클래스를 가진 요소를 포함해야 합니다.

```js
const petNameEl = document.querySelector('.pet-name');
assert.exists(petNameEl);
```

두 번째 뷰는 제공된 펫 이름을 포함해야 합니다.

```js
const petNameEl = document.querySelector('.pet-name');
assert.strictEqual(petNameEl.textContent, _PET_NAME);
```

두 번째 뷰는 세 개의 `button`를 포함해야 합니다.

```js
const buttons = document.querySelectorAll('button');
assert.equal(buttons.length, 3);
```

정확히 하나의 버튼이 `#eat-action` 아이디를 가져야 합니다.

```js
const eatButton = document.querySelectorAll('button#eat-action');
assert.equal(eatButton.length, 1);
```

정확히 하나의 버튼이 `#play-action` 아이디를 가져야 합니다.

```js
const playButton = document.querySelectorAll('button#play-action');
assert.equal(playButton.length, 1);
```

정확히 하나의 버튼이 `#sleep-action` 아이디를 가져야 합니다.

```js
const sleepButton = document.querySelectorAll('button#sleep-action');
assert.equal(sleepButton.length, 1);
```

정확히 세 개의 `.stat` 클래스를 가진 요소가 있어야 합니다.

```js
const statEls = document.querySelectorAll('.stat');
assert.equal(statEls.length, 3);
```

각 펫 상태는 정확히 하나의 `.stat-value` 클래스를 가진 요소를 포함해야 합니다.

```js
const statEls = document.querySelectorAll('.stat');
for (const statEl of statEls) {
  assert.equal(statEl.querySelectorAll('.stat-value').length, 1);
}
```

각 `.stat-value` 요소는 하나의 숫자를 포함해야 합니다.

```js
const statValueEls = document.querySelectorAll('.stat-value');
for (const statEl of statValueEls) {
  assert.match(statEl.textContent, /-?\d+\.*\d*/);
}
```

각 `.stat-value` 숫자는 `[0,100]` 범위 내에 있어야 하며, 범위는 포함합니다.

```js
const statValueEls = document.querySelectorAll('.stat-value');
for (const statEl of statValueEls) {
  const statValue = parseFloat(statEl.textContent);
  assert.isAtMost(statValue, 100);
  assert.isAtLeast(statValue, 0);
}
```

정확히 하나의 `.stat` 요소가 `Hunger` 텍스트를 포함해야 하며, 대소문자를 구분하지 않습니다.

```js
const matches = Array.from(document.querySelectorAll('.stat')).filter(el => /hunger/i.test(el.textContent));
assert.equal(matches.length, 1);
```

하나의 `.stat` 요소가 `Energy` 텍스트를 포함해야 하며, 대소문자를 구분하지 않습니다.

```js
const matches = Array.from(document.querySelectorAll('.stat')).filter(el => /energy/i.test(el.textContent));
assert.equal(matches.length, 1);
```

하나의 `.stat` 요소가 `Happiness` 텍스트를 포함해야 하며, 대소문자를 구분하지 않습니다.

```js
const matches = Array.from(document.querySelectorAll('.stat')).filter(el => /happiness/i.test(el.textContent));
assert.equal(matches.length, 1);
```

`Hunger` 상태는 0에서 시작해야 합니다.

```js
const hungerStatEl = Array.from(document.querySelectorAll('.stat')).find(el => /hunger/i.test(el.textContent));
const hungerValue = parseFloat(hungerStatEl.querySelector('.stat-value').textContent);
assert.equal(hungerValue, 0);
```

`Energy` 상태는 100에서 시작해야 합니다.

```js
const energyStatEl = Array.from(document.querySelectorAll('.stat')).find(el => /energy/i.test(el.textContent));
const energyValue = parseFloat(energyStatEl.querySelector('.stat-value').textContent);
assert.equal(energyValue, 100);
```

`Happiness` 상태는 100에서 시작해야 합니다.

```js
const happinessStatEl = Array.from(document.querySelectorAll('.stat')).find(el => /happiness/i.test(el.textContent));
const happinessValue = parseFloat(happinessStatEl.querySelector('.stat-value').textContent);
assert.equal(happinessValue, 100);
```

`Eat` 버튼을 클릭하면 `Hunger` 값이 줄고 `Energy` 값이 늘어납니다. 만약 `Hunger`가 0이거나 `Energy`가 100이면 값이 변하지 않아야 합니다.

```js
const statEls = () => Array.from(document.querySelectorAll('.stat'));
const hungerBefore = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyBefore = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
document.querySelector('#eat-action').click();
_clock.tick(50);
const hungerAfter = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyAfter = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
if (hungerBefore !== 0) assert.isBelow(hungerAfter, hungerBefore);
else assert.equal(hungerAfter, hungerBefore);
if (energyBefore !== 100) assert.isAbove(energyAfter, energyBefore);
else assert.equal(energyAfter, energyBefore);
```

`Play` 버튼을 클릭하면 `Energy` 값이 줄고 `Happiness` 값이 늘어납니다. 만약 `Energy`가 0이거나 `Happiness`가 100이면 값이 변하지 않아야 합니다.

```js
const statEls = () => Array.from(document.querySelectorAll('.stat'));
const energyBefore = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
const happinessBefore = parseFloat(statEls().find(el => /happiness/i.test(el.textContent)).querySelector('.stat-value').textContent);
document.querySelector('#play-action').click();
_clock.tick(50);
const energyAfter = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
const happinessAfter = parseFloat(statEls().find(el => /happiness/i.test(el.textContent)).querySelector('.stat-value').textContent);
if (energyBefore !== 0) assert.isBelow(energyAfter, energyBefore);
else assert.equal(energyAfter, energyBefore);
if (happinessBefore !== 100) assert.isAbove(happinessAfter, happinessBefore);
else assert.equal(happinessAfter, happinessBefore);
```

`Sleep` 버튼을 클릭하면 `Hunger` 값이 늘고 `Energy` 값이 늘어납니다. 만약 `Hunger`가 100이거나 `Energy`가 100이면 값이 변하지 않아야 합니다.

```js
const statEls = () => Array.from(document.querySelectorAll('.stat'));
const hungerBefore = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyBefore = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
document.querySelector('#sleep-action').click();
_clock.tick(50);
const hungerAfter = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyAfter = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
if (hungerBefore !== 100) assert.isAbove(hungerAfter, hungerBefore);
else assert.equal(hungerAfter, hungerBefore);
if (energyBefore !== 100) assert.isAbove(energyAfter, energyBefore);
else assert.equal(energyAfter, energyBefore);
```

과도한 유휴 시간 후 `Hunger` 상태는 100이어야 합니다.

```js
await _clock.tickAsync(1_200_000);

const hungerStatEl = Array.from(document.querySelectorAll('.stat')).find(el => /hunger/i.test(el.textContent));
const hungerValue = parseFloat(hungerStatEl.querySelector('.stat-value').textContent);
assert.equal(hungerValue, 100);
```

과도한 유휴 시간 후 `Energy` 상태는 100이어야 합니다.

```js
const energyStatEl = Array.from(document.querySelectorAll('.stat')).find(el => /energy/i.test(el.textContent));
const energyValue = parseFloat(energyStatEl.querySelector('.stat-value').textContent);
assert.equal(energyValue, 100);
```

과도한 유휴 시간 후 `Happiness` 상태는 0이어야 합니다.

```js
const happinessStatEl = Array.from(document.querySelectorAll('.stat')).find(el => /happiness/i.test(el.textContent));
const happinessValue = parseFloat(happinessStatEl.querySelector('.stat-value').textContent);
assert.equal(happinessValue, 0);
```

`Eat` 버튼을 클릭하면 `Hunger` 값이 줄고 `Energy` 값이 늘어납니다. 만약 `Hunger`가 0이거나 `Energy`가 100이면 값이 변하지 않아야 합니다.

```js
const statEls = () => Array.from(document.querySelectorAll('.stat'));
const hungerBefore = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyBefore = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
document.querySelector('#eat-action').click();
_clock.tick(50);
const hungerAfter = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyAfter = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
if (hungerBefore !== 0) assert.isBelow(hungerAfter, hungerBefore);
else assert.equal(hungerAfter, hungerBefore);
if (energyBefore !== 100) assert.isAbove(energyAfter, energyBefore);
else assert.equal(energyAfter, energyBefore);
```

`Play` 버튼을 클릭하면 `Energy` 값이 줄고 `Happiness` 값이 늘어납니다. 만약 `Energy`가 0이거나 `Happiness`가 100이면 값이 변하지 않아야 합니다.

```js
const statEls = () => Array.from(document.querySelectorAll('.stat'));
const energyBefore = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
const happinessBefore = parseFloat(statEls().find(el => /happiness/i.test(el.textContent)).querySelector('.stat-value').textContent);
document.querySelector('#play-action').click();
_clock.tick(50);
const energyAfter = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
const happinessAfter = parseFloat(statEls().find(el => /happiness/i.test(el.textContent)).querySelector('.stat-value').textContent);
if (energyBefore !== 0) assert.isBelow(energyAfter, energyBefore);
else assert.equal(energyAfter, energyBefore);
if (happinessBefore !== 100) assert.isAbove(happinessAfter, happinessBefore);
else assert.equal(happinessAfter, happinessBefore);
```

`Sleep` 버튼을 클릭하면 `Hunger` 값이 늘고 `Energy` 값이 늘어납니다. 만약 `Hunger`가 100이거나 `Energy`가 100이면 값이 변하지 않아야 합니다.

```js
const statEls = () => Array.from(document.querySelectorAll('.stat'));
const hungerBefore = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyBefore = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
document.querySelector('#sleep-action').click();
_clock.tick(50);
const hungerAfter = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyAfter = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
if (hungerBefore !== 100) assert.isAbove(hungerAfter, hungerBefore);
else assert.equal(hungerAfter, hungerBefore);
if (energyBefore !== 100) assert.isAbove(energyAfter, energyBefore);
else assert.equal(energyAfter, energyBefore);
```

`PetMood` 열거자가 있어야 합니다.

```js
assert.match(code, /\s*enum PetMood\s*/);
```

`PetMood` 열거자는 `HAPPY`를 포함해야 합니다.

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*HAPPY[\s*\w+,\s*]*}/);
```

`PetMood` 열거자는 `EXCITED`를 포함해야 합니다.

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*EXCITED[\s*\w+,\s*]*}/);
```

`PetMood` 열거자는 `CONTENT`를 포함해야 합니다.

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*CONTENT[\s*\w+,\s*]*}/);
```

`PetMood` 열거자는 `SAD`'를 포함해야 합니다.

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*SAD[\s*\w+,\s*]*}/);
```

`PetMood` 열거자는 `TIRED`를 포함해야 합니다.

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*TIRED[\s*\w+,\s*]*}/);
```

`PetMood` 열거자는 `SICK`를 포함해야 합니다.

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*SICK[\s*\w+,\s*]*}/);
```

`PetMood` 열거자는 `HUNGRY`를 포함해야 합니다.

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*HUNGRY[\s*\w+,\s*]*}/);
```

각 `Record<PetMood, string`에 값을 매핑하는 `PetMood`이 있어야 합니다.

```js
assert.match(code, /\s*(?:const|let|var)\s+\w+:\s+Record<PetMood,\s*string>/);
```

`PetMood.HAPPY`에 매핑된 값이 있어야 합니다.

```js
assert.match(code, /\[PetMood\.HAPPY\]\s*:/);
```

`PetMood.EXCITED`에 매핑된 값이 있어야 합니다.

```js
assert.match(code, /\[PetMood\.EXCITED\]\s*:/);
```

`PetMood.CONTENT`에 매핑된 값이 있어야 합니다.

```js
assert.match(code, /\[PetMood\.CONTENT\]\s*:/);
```

`PetMood.SAD`에 매핑된 값이 있어야 합니다.

```js
assert.match(code, /\[PetMood\.SAD\]\s*:/);
```

`PetMood.TIRED`에 매핑된 값이 있어야 합니다.

```js
assert.match(code, /\[PetMood\.TIRED\]\s*:/);
```

`PetMood.SICK`에 매핑된 값이 있어야 합니다.

```js
assert.match(code, /\[PetMood\.SICK\]\s*:/);
```

`PetMood.HUNGRY`에 매핑된 값이 있어야 합니다.

```js
assert.match(code, /\[PetMood\.HUNGRY\]\s*:/);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="utf-8"/>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Digital Pet Game</title>

    <link rel="preconnect" href="https://fonts.googleapis.com">
    <link rel="preconnect" href="https://fonts.gstatic.com" crossorigin>
    <link href="https://fonts.googleapis.com/css2?family=Tektur:wght@400..900&display=swap" rel="stylesheet">

        <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.tsx"
    ></script>
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { PetGame } from './index.tsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<PetGame />);
    </script>
  </body>
</html>
```

```css
:root {
  /* Backgrounds */
  --bg-gradient-start: #dbeafe;
  --bg-gradient-mid: #faf5ff;
  --bg-gradient-end: #fce7f3;

  /* Text */
  --text-primary: #1f2937;
  --text-muted: #6b7280;
  --text-muted-2: #9ca3af;

  /* Card / surface */
  --card-bg-translucent: rgba(255, 255, 255, 0.6);
  --card-bg-solid: rgba(255, 255, 255, 0.8);

  /* Accent / brand */
  --accent-purple: #8b5cf6;
  --accent-purple-dark: #7c3aed;

  /* Actions */
  --action-feed: #10b981;
  --action-feed-dark: #059669;
  --action-play: #3b82f6;
  --action-play-dark: #2563eb;
  --action-sleep: #8b5cf6;
  --action-sleep-dark: #7c3aed;
  --action-heal: #ec4899;
  --action-heal-dark: #db2777;

  /* Stats */
  --stat-high: #10b981;
  --stat-medium: #f59e0b;
  --stat-low: #ef4444;

  /* UI greys */
  --ui-bg-light: #e5e7eb;
  --ui-fill: #8b5cf6; /* used for exp fill as accent */

  /* Shadows */
  --shadow-strong: rgba(0, 0, 0, 0.25);
  --shadow-soft: rgba(0, 0, 0, 0.1);
}

html {
  /*color-scheme: light dark;*/
  font-family: system-ui;
  display: flex;
  justify-content: center;
  font-family:
    -apple-system, BlinkMacSystemFont, "Segoe UI", "Roboto", "Oxygen", "Ubuntu",
    "Cantarell", "Fira Sans", "Droid Sans", "Helvetica Neue", sans-serif;
  -webkit-font-smoothing: antialiased;
  -moz-osx-font-smoothing: grayscale;
  background: linear-gradient(135deg, #dbeafe 0%, #faf5ff 50%, #fce7f3 100%);
  min-height: 100vh;
  padding: 1rem;
}

header {
  text-align: center;
  margin-bottom: 2rem;
}

body {
  max-width: 400px;
}

.base-container {
  background: rgba(255, 255, 255, 0.6);
  -webkit-backdrop-filter: blur(10px);
  backdrop-filter: blur(10px);
  border-radius: 1.5rem;
  padding: 2rem;
  box-shadow: 0 25px 50px -12px rgba(0, 0, 0, 0.25);
  .pet-name {
    text-align: center;
  }
}

.game-container {
  margin-bottom: 1.5rem;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
}

/* Stats Grid */
.stats-grid {
  display: grid;
  grid-template-columns: repeat(2, 1fr);
  gap: 1rem;
  margin-bottom: 1.5rem;
}

.stat-bar {
  background: rgba(255, 255, 255, 0.8);
  -webkit-backdrop-filter: blur(10px);
  backdrop-filter: blur(10px);
  border-radius: 0.75rem;
  padding: 1rem;
  box-shadow: 0 1px 3px 0 rgba(0, 0, 0, 0.1);
  transition: all 0.2s ease;
}

.stat-bar:hover {
  box-shadow: 0 4px 6px -1px rgba(0, 0, 0, 0.1);
}

.stat-header {
  display: flex;
  align-items: center;
  justify-content: space-between;
  margin-bottom: 0.5rem;
}

.stat-label {
  display: flex;
  align-items: center;
  gap: 0.5rem;
}

.stat-icon {
  font-size: 1.25rem;
}

.stat-name {
  font-weight: 500;
  color: #374151;
}

.stat-value {
  font-weight: bold;
  color: #1f2937;
}

.stat-progress {
  width: 100%;
  height: 0.75rem;
  background: #e5e7eb;
  border-radius: 9999px;
  overflow: hidden;
}

.stat-fill {
  height: 100%;
  transition: width 0.5s ease-out;
  border-radius: 9999px;
}

.stat-fill.high {
  background: #10b981;
}

.stat-fill.medium {
  background: #f59e0b;
}

.stat-fill.low {
  background: #ef4444;
}

.pet-buttons {
  width: auto;
  display: flex;
  flex-direction: row;
  gap: 1rem;
}

.pet-buttons .pet-button {
  min-width: 80px;
  background: #c2214a;
  color: white;
  border-radius: 13%;
  border-style: solid;
  border-color: #870b04;
  box-shadow: rgba(120, 25, 23, 0.2) 2px 2px 4px 2px;
  &:active {
    transform: scale(0.95);
    box-shadow: 0 2px 4px rgba(0, 0, 0, 0.15) inset;
  }
  text-align: center;
  cursor: pointer;
  padding: 0.5rem 1rem;
}

.pet-sprite {
  font-size: xxx-large;
  text-align: center;
}

.data-management * {
  margin: 0 0.5rem 0 0;
}

.info-panel {
  text-align: center;
}
.info-panel #hud,
.info-panel .start-questions {
  display: flex;
  align-items: center;
  flex-direction: column;
  gap: 1rem;
}

.info-panel #hud p {
  margin: 0rem;
}

.info-panel .start-questions div * {
  margin: 0 0.5rem 0 0;
}

#pet-fact {
  color: var(--text-muted-2);
  margin: 0 0 0.5rem 0;
}
#set-name-btn {
  display: inline-block;
  padding: 14px 28px;
  font-size: 1.15rem;
  font-weight: 600;
  background: linear-gradient(135deg, #6b8cff, #5ec2ff);
  color: #fff;
  border: none;
  border-radius: 10px;
  box-shadow: 0 6px 18px rgba(92, 132, 255, 0.25);
  cursor: pointer;
  transition:
    transform 0.08s ease,
    box-shadow 0.12s ease,
    opacity 0.12s;
}

#set-name-btn:hover {
  opacity: 0.98;
}

#set-name-btn:active,
#set-name-btn.pressed {
  transform: translateY(2px) scale(0.995);
  box-shadow: 0 3px 8px rgba(92, 132, 255, 0.18);
}

#set-name-btn:focus {
  outline: 3px solid rgba(94, 190, 255, 0.28);
  outline-offset: 2px;
}
```

```tsx
const { useState } = React;

export const PetGame = () => {};
```

# --solutions--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="utf-8"/>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Digital Pet Game</title>

    <link rel="preconnect" href="https://fonts.googleapis.com">
    <link rel="preconnect" href="https://fonts.gstatic.com" crossorigin>
    <link href="https://fonts.googleapis.com/css2?family=Tektur:wght@400..900&display=swap" rel="stylesheet">
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.tsx"
    ></script>
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { PetGame } from './index.tsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<PetGame />);
    </script>
  </body>
</html>
```

```css
:root {
  /* Backgrounds */
  --bg-gradient-start: #dbeafe;
  --bg-gradient-mid: #faf5ff;
  --bg-gradient-end: #fce7f3;

  /* Text */
  --text-primary: #1f2937;
  --text-muted: #6b7280;
  --text-muted-2: #9ca3af;

  /* Card / surface */
  --card-bg-translucent: rgba(255, 255, 255, 0.6);
  --card-bg-solid: rgba(255, 255, 255, 0.8);

  /* Accent / brand */
  --accent-purple: #8b5cf6;
  --accent-purple-dark: #7c3aed;

  /* Actions */
  --action-feed: #10b981;
  --action-feed-dark: #059669;
  --action-play: #3b82f6;
  --action-play-dark: #2563eb;
  --action-sleep: #8b5cf6;
  --action-sleep-dark: #7c3aed;
  --action-heal: #ec4899;
  --action-heal-dark: #db2777;

  /* Stats */
  --stat-high: #10b981;
  --stat-medium: #f59e0b;
  --stat-low: #ef4444;

  /* UI greys */
  --ui-bg-light: #e5e7eb;
  --ui-fill: #8b5cf6; /* used for exp fill as accent */

  /* Shadows */
  --shadow-strong: rgba(0, 0, 0, 0.25);
  --shadow-soft: rgba(0, 0, 0, 0.1);
}

html {
  /*color-scheme: light dark;*/
  font-family: system-ui;
  display: flex;
  justify-content: center;
  font-family:
    -apple-system, BlinkMacSystemFont, "Segoe UI", "Roboto", "Oxygen", "Ubuntu",
    "Cantarell", "Fira Sans", "Droid Sans", "Helvetica Neue", sans-serif;
  -webkit-font-smoothing: antialiased;
  -moz-osx-font-smoothing: grayscale;
  background: linear-gradient(135deg, #dbeafe 0%, #faf5ff 50%, #fce7f3 100%);
  min-height: 100vh;
  padding: 1rem;
}

header {
  text-align: center;
  margin-bottom: 2rem;
}

body {
  max-width: 400px;
}

.base-container {
  background: rgba(255, 255, 255, 0.6);
  -webkit-backdrop-filter: blur(10px);
  backdrop-filter: blur(10px);
  border-radius: 1.5rem;
  padding: 2rem;
  box-shadow: 0 25px 50px -12px rgba(0, 0, 0, 0.25);
  .pet-name {
    text-align: center;
  }
}

.game-container {
  margin-bottom: 1.5rem;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
}

/* Stats Grid */
.stats-grid {
  display: grid;
  grid-template-columns: repeat(2, 1fr);
  gap: 1rem;
  margin-bottom: 1.5rem;
}

.stat-bar {
  background: rgba(255, 255, 255, 0.8);
  -webkit-backdrop-filter: blur(10px);
  backdrop-filter: blur(10px);
  border-radius: 0.75rem;
  padding: 1rem;
  box-shadow: 0 1px 3px 0 rgba(0, 0, 0, 0.1);
  transition: all 0.2s ease;
}

.stat-bar:hover {
  box-shadow: 0 4px 6px -1px rgba(0, 0, 0, 0.1);
}

.stat-header {
  display: flex;
  align-items: center;
  justify-content: space-between;
  margin-bottom: 0.5rem;
}

.stat-label {
  display: flex;
  align-items: center;
  gap: 0.5rem;
}

.stat-icon {
  font-size: 1.25rem;
}

.stat-name {
  font-weight: 500;
  color: #374151;
}

.stat-value {
  font-weight: bold;
  color: #1f2937;
}

.stat-progress {
  width: 100%;
  height: 0.75rem;
  background: #e5e7eb;
  border-radius: 9999px;
  overflow: hidden;
}

.stat-fill {
  height: 100%;
  transition: width 0.5s ease-out;
  border-radius: 9999px;
}

.stat-fill.high {
  background: #10b981;
}

.stat-fill.medium {
  background: #f59e0b;
}

.stat-fill.low {
  background: #ef4444;
}

.pet-buttons {
  width: auto;
  display: flex;
  flex-direction: row;
  gap: 1rem;
}

.pet-buttons .pet-button {
  min-width: 80px;
  background: #c2214a;
  color: white;
  border-radius: 13%;
  border-style: solid;
  border-color: #870b04;
  box-shadow: rgba(120, 25, 23, 0.2) 2px 2px 4px 2px;
  &:active {
    transform: scale(0.95);
    box-shadow: 0 2px 4px rgba(0, 0, 0, 0.15) inset;
  }
  text-align: center;
  cursor: pointer;
  padding: 0.5rem 1rem;
}

.pet-sprite {
  font-size: xxx-large;
  text-align: center;
}

.data-management * {
  margin: 0 0.5rem 0 0;
}

.info-panel {
  text-align: center;
}
.info-panel #hud,
.info-panel .start-questions {
  display: flex;
  align-items: center;
  flex-direction: column;
  gap: 1rem;
}

.info-panel #hud p {
  margin: 0rem;
}

.info-panel .start-questions div * {
  margin: 0 0.5rem 0 0;
}

#pet-fact {
  color: var(--text-muted-2);
  margin: 0 0 0.5rem 0;
}
#set-name-btn {
  display: inline-block;
  padding: 14px 28px;
  font-size: 1.15rem;
  font-weight: 600;
  background: linear-gradient(135deg, #6b8cff, #5ec2ff);
  color: #fff;
  border: none;
  border-radius: 10px;
  box-shadow: 0 6px 18px rgba(92, 132, 255, 0.25);
  cursor: pointer;
  transition:
    transform 0.08s ease,
    box-shadow 0.12s ease,
    opacity 0.12s;
}

#set-name-btn:hover {
  opacity: 0.98;
}

#set-name-btn:active,
#set-name-btn.pressed {
  transform: translateY(2px) scale(0.995);
  box-shadow: 0 3px 8px rgba(92, 132, 255, 0.18);
}

#set-name-btn:focus {
  outline: 3px solid rgba(94, 190, 255, 0.28);
  outline-offset: 2px;
}
```

```tsx
const { useState, useEffect } = React;

interface StatBarProps {
  label: string;
  value: number;
  icon: string;
  reverse?: boolean;
}

function getStatColorClass(value: number): string {
  if (value >= 70) return "high";
  if (value >= 40) return "medium";
  return "low";
}

function StatBar({ label, value, icon, reverse = false }: StatBarProps) {
  const colorClass = getStatColorClass(!reverse ? value : 100 - value);

  return (
    <div className="stat-bar stat">
      <div className="stat-header">
        <div className="stat-label">
          <span className="stat-icon">{icon}</span>
          <span className="stat-name">{label}</span>
        </div>
        <span className="stat-value">{Math.round(value)}%</span>
      </div>

      <div className="stat-progress">
        <div
          className={`stat-fill ${colorClass}`}
          style={{ width: `${value}%` }}
        />
      </div>
    </div>
  );
}

enum PetAction {
  NONE,
  EAT,
  PLAY,
  SLEEP,
}

interface Pet {
  name: string;
  happiness: number;
  hunger: number;
  energy: number;
  species: string;
}

enum PetMood {
  HAPPY,
  EXCITED,
  CONTENT,
  SAD,
  TIRED,
  SICK,
  HUNGRY,
}

export const STAT_DECAY_RATES = {
  hunger: 10,
  happiness: 5,
  energy: 5,
};

const UPDATE_INTERVAL = 30 * 1000; // 30 seconds

export function calculatePetMood(pt: Pet): PetMood {
  const { hunger, happiness, energy } = pt;

  if (hunger > 70) return PetMood.HUNGRY;
  if (energy < 30) return PetMood.TIRED;
  if (happiness < 30) return PetMood.SAD;
  if (happiness > 80 && energy > 70) return PetMood.EXCITED;
  if (happiness > 60) return PetMood.HAPPY;

  return PetMood.CONTENT;
}

const moodEmojiMap: Record<PetMood, string> = {
  [PetMood.HAPPY]: "😺", // grinning cat
  [PetMood.EXCITED]: "😻", // heart-eyes cat
  [PetMood.CONTENT]: "😸", // smiling cat
  [PetMood.SAD]: "😿", // crying cat
  [PetMood.TIRED]: "😽", // kissing cat (closest to sleepy)
  [PetMood.SICK]: "🙀", // weary cat (looks unwell)
  [PetMood.HUNGRY]: "😹", // joy cat (closest to hungry/funny)
};

interface usePetProps {
  isGameStarted: boolean;
  setGameStarted: React.Dispatch<React.SetStateAction<boolean>>;
}
// hook where all pet logic is handled
function usePet({ isGameStarted }: usePetProps) {
  let [pet, setPet] = useState<Pet>({
    name: "",
    happiness: 100,
    hunger: 0,
    energy: 100,
    species: "Cat",
  });
  // Auto-decay stats over time
  useEffect(() => {
    if (!isGameStarted) return;

    const interval = setInterval(() => {
      setPet((pet) => ({
        ...pet,
        happiness: Math.max(pet.happiness - 5, 0),
        hunger: Math.min(pet.hunger + 5, 100),
      }));
    }, UPDATE_INTERVAL);

    return () => clearInterval(interval);
  }, [isGameStarted]);

  function doAction(action: PetAction) {
    switch (action) {
      case PetAction.EAT:
        feedPet();
        break;
      case PetAction.SLEEP:
        restPet();
        break;
      case PetAction.PLAY:
        playWithPet();
        break;
    }
  }

  function feedPet() {
    setPet((pet) => ({
      ...pet,
      hunger: Math.max(pet.hunger - STAT_DECAY_RATES.hunger, 0),
      energy: Math.min(pet.energy + STAT_DECAY_RATES.energy, 100),
    }));
  }

  function playWithPet() {
    setPet((pet) => ({
      ...pet,
      energy: Math.max(pet.energy - STAT_DECAY_RATES.energy, 0),
      happiness: Math.min(pet.happiness + STAT_DECAY_RATES.happiness, 100),
    }));
  }

  function restPet() {
    setPet((pet) => ({
      ...pet,
      hunger: Math.min(pet.hunger + STAT_DECAY_RATES.hunger, 100),
      energy: Math.min(pet.energy + STAT_DECAY_RATES.energy, 100),
    }));
  }

  return {
    pet,
    doAction,
    setName: (name: string) => setPet({ ...pet, name }),
  };
}

export function PetGame() {
  const [isGameStarted, setGameStarted] = useState(false);

  let { pet, doAction, setName } = usePet({
    isGameStarted,
    setGameStarted,
  });

  const [fact, setFact] = useState<string | null>("");

  useEffect(() => {
    fetch("https://cat-facts-api.freecodecamp.rocks/api/catfacts/random").then(
      res => res.json()
    ).then(
      fact => setFact(fact)
    ).catch(err => {
      setFact("Sorry, we're not able to retrieve your cat fact right now!");
    });
  }, []);

  function startGame() {
    const petName = (document.getElementById("pet-name") as HTMLInputElement)
      .value;
    if (!petName) return; //do not let user to start game if no name selected
    setName(petName);
    setGameStarted(true);
  }

  return (
    <main>
      <header>
        <h1>Digital Pet Game</h1>
        <p>Take care of your virtual companion!</p>
      </header>

      {isGameStarted && (
        <section className="base-container game-container">
          <div className="pet-screen">
            <div className="pet-sprite">
              {moodEmojiMap[calculatePetMood(pet)]}
            </div>
            <h2 className="pet-name">{pet.name}</h2>
          </div>
          <div className="pet-buttons">
            <button
              onClick={() => doAction(PetAction.EAT)}
              className="pet-button pet-buttons-left"
              id="eat-action"
            >
              EAT
            </button>
            <button
              onClick={() => doAction(PetAction.PLAY)}
              className="pet-button pet-buttons-center"
              id="play-action"
            >
              PLAY
            </button>
            <button
              onClick={() => doAction(PetAction.SLEEP)}
              className="pet-button pet-buttons-right"
              id="sleep-action"
            >
              SLEEP
            </button>
          </div>
        </section>
      )}

      {isGameStarted && (
        <section className="stats-grid">
          <StatBar label="Hunger" value={pet.hunger} icon="🍽️" reverse />
          <StatBar label="Happiness" value={pet.happiness} icon="😊" />
          <StatBar label="Energy" value={pet.energy} icon="⚡" />
        </section>
      )}

      <section className="base-container info-panel">
        {!isGameStarted ? (
          <form className="start-questions">
            <label htmlFor="pet-name">What is your pet's name?</label>
            <input
              id="pet-name"
              name="pet-name"
              required={true}
              pattern="[A-Za-z0-9]{1,20}"
            />

            <button id="set-name-btn" onClick={startGame}>
              Start Game
            </button>
          </form>
        ) : (
          <div id="hud">
            <p id="pet-species">Species: {pet.species}</p>
            <p id="pet-fact">
              <b>Pet Fact:</b> {fact}
            </p>
          </div>
        )}
      </section>
    </main>
  );
}
```
