---
id: 5dc17d3bf86c76b9248c6eb4
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Crie um elemento `p` abaixo do elemento `h2` e dê a ele o seguinte texto:

`Everyone loves cute cats online!`

# --hints--

O elemento `p` deve ter uma tag de abertura. As tags de abertura possuem a seguinte sintaxe: `<elementName>`.

```js
assert.exists(document.querySelector('p'));
```

Suas tags `p` devem estar em letras minúsculas. Por convenção, todas as tags HTML são escritas em letras minúsculas.

```js
assert.notMatch(code, /<\/?P>/);
```

O elemento `p` deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert.match(code, /<\/p\>/);
```

O texto do seu elemento `p` deve ser `Todos adoram gatos fofos online!` Você ou omitiu o texto ou cometeu um erro de digitação.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  ?.innerText?.trim().replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

O elemento `p` deve estar abaixo do elemento `h2`. Eles estão na ordem errada.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H2'), collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
