---
id: 5dfa3589eacea3f48c6300ae
title: Passo 20
challengeType: 0
dashedName: step-20
---

# --description--

Dentro do segundo elemento `section`, adicione um novo elemento `h2` com o texto `Cat Lists`.

# --hints--

O elemento `section` deve ter uma tag de abertura.As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
assert.lengthOf(code.match(/<\/section>/g), 2);
```

O elemento `h2` deve ter uma tag de abertura.As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('h2'), 2);
```

O elemento `h2` deve ter uma tag de fechamento.As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert.lengthOf(code.match(/<\/h2\>/g), 2);
```

O segundo elemento `h2` deve estar logo acima da tag de fechamento do segundo elemento `section`.Ele não está na posição correta.

```js
const secondSection = document.querySelectorAll('section')[1];
assert.equal(secondSection?.lastElementChild?.nodeName, 'H2');
```

O segundo elemento `h2` deve ter o texto `Cat Lists`.Você omitiu o texto ou tem um erro de digitação.

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    ?.lastElementChild?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>
        
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
