---
id: 5dfb655eeacea3f48c6300b3
title: Passo 25
challengeType: 0
dashedName: step-25
---

# --description--

O elemento `figure` representa um conteúdo auto-contido e permitirá que você associe uma imagem com uma legenda.

Aninhe a imagem que você acabou de adicionar dentro de um elemento `figure`.

# --hints--

O elemento `figure` deve ter uma tag de abertura. As tags de abertura possuem a seguinte sintaxe: `<elementName>`.

```js
assert.exists(document.querySelector('figure'));
```

O elemento `figure` deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert.match(code, /<\/figure\>/);
```

Deve haver um elemento `figure` logo acima da tag de fechamento do segundo elemento `section`.

```js
assert.equal(document.querySelectorAll('section')[1]?.lastElementChild.nodeName, 'FIGURE');
```

O elemento `img` da lasanha deve estar dentro do elemento `figure`.

```js
assert.exists(document.querySelector('figure > img'));
assert.equal(
    document.querySelector('figure > img').getAttribute('src').toLowerCase(),
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        
        <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
        
--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```
