---
id: 5ef9b03c81a63668521804d1
title: Passo 28
challengeType: 0
dashedName: step-28
---

# --description--

Após o elemento `figure`, adicione outro elemento `h3` com o texto:

`Top 3 things cats hate:`

# --hints--

O elemento `h3` deve estar logo acima da tag de fechamento de fechamento do segundo elemento `section`. Certifique-se de que ele tenha uma tag de abertura e fechamento.

```js
assert.equal(
  document.querySelectorAll('main > section')[1]?.lastElementChild.nodeName,
    'H3'
);
assert.lengthOf(code.match(/<\/h3\>/g),  2);
```

O novo elemento `h3` deve ter o texto `Top 3 things cats hate:`. Certifique-se de incluir os dois pontos ao final do texto.

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    ?.lastElementChild.innerText?.trim().toLowerCase()
    .replace(/\s+/g, ' '), 'top 3 things cats hate:'
);
```

Deve haver um elemento `figure` acima do novo elemento `h3`. Você pode ter excluído o elemento `figure` acidentalmente.

```js
const secondSectionLastElemNode = document.querySelectorAll('main > section')[1]
  ?.lastElementChild;
assert.equal(secondSectionLastElemNode?.nodeName, 'H3');
assert.equal(secondSectionLastElemNode.previousElementSibling.nodeName, 'FIGURE');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
