---
id: 5ef9b03c81a63668521804eb
title: Passo 39
challengeType: 0
dashedName: step-39
---

# --description--

O elemento `title` determina o que os navegadores mostrarão na barra de título ou na aba da página.

Adicione um elemento `title` dentro do elemento `head` usando o texto abaixo: 

`CatPhotoApp`

# --hints--

Você excluiu o elemento `head` ou está faltando nele uma tag de abertura ou de fechamento.

```js
assert.match(code, /\<head\>/);
assert.match(code, /\<\/head\>/);
```

Seu elemento `title` deve estar aninhado no elemento `head`. Certifique-se de adicionar uma tag de abertura e uma tag de fechamento para o elemento `title`.

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<head\>\<title\>.*\<\/title\>\<\/head\>/);
```

O elemento `title` deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert.match(code, /\<\/title\>/);
```

O elemento `title` deve conter o texto `CatPhotoApp`. Você omitiu o texto ou tem um erro de digitação.

```js
assert.equal(document.title?.toLowerCase(), 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <head>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
