---
id: 5ef9b03c81a63668521804e8
title: Passo 36
challengeType: 0
dashedName: step-36
---

# --description--

Coloque um elemento `p` com o texto `No Copyright - freeCodeCamp.org` dentro do elemento `footer`.

# --hints--

Você excluiu o elemento `footer` ou está faltando nele uma tag de abertura ou de fechamento.

```js
assert.exists(document.querySelector('footer'));
assert.match(code, /<\/footer>/);
```

O elemento `footer` deve ter um elemento `p`.Certifique-se de adicionar uma tag de abertura e uma tag de fechamento para o elemento `p`.

```js
assert.exists(document.querySelector('footer > p'));
```

O elemento `footer` deve ter uma tag de fechamento.As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
const pElemClosingTags = code.match(/<\/p\>/g);
assert.exists(pElemClosingTags);
assert.lengthOf(pElemClosingTags, 3);
```

O texto do elemento `p` deve ser `No Copyright - freeCodeCamp.org`.Você omitiu o texto, digitou o texto incorretamente ou ele não está entre as tags de abertura e fechamento do elemento `p`.

```js
const extraSpacesRemoved = document.querySelectorAll('footer > p')[0]?.innerText?.trim().replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /^No Copyright - freeCodeCamp\.org$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
--fcc-editable-region--
    <footer>
      
    </footer>
--fcc-editable-region--
  </body>
</html>
```
