---
id: 5ef9b03c81a63668521804ee
title: Passo 41
challengeType: 0
dashedName: step-41
---

# --description--

Todas as páginas devem começar com `<!DOCTYPE html>`. Essa string especial é conhecida como uma <dfn>declaração</dfn> e garante que o navegador tentará atender às especificações gerais do setor.

`<!DOCTYPE html>` diz aos navegadores que o documento é um documento HTML5, que é a versão mais recente do HTML.

Adicione essa declaração como a primeira linha do código.

# --hints--

O código deve começar com a declaração `<!DOCTYPE html>`. Você pode ter omitido a declaração, ter um erro de digitação ou a declaração pode não ser a primeira linha do código.

```js
assert.match(code, /<\s*!doctype\s+html\s*>/i);
```

`<!DOCTYPE html>` deve estar localizada na parte superior do documento.

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /^<\!doctypehtml><html/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
<html lang="en">
  <head>
    <title>CatPhotoApp</title>
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
