---
id: 5ef9b03c81a63668521804d2
title: Passo 29
challengeType: 0
dashedName: step-29
---

# --description--

O código de uma lista ordenada (`ol`) é semelhante ao de uma lista não ordenada, mas a lista de itens em uma lista ordenada é numerada quando eles são exibidos.

Abaixo do elemento `h3` , adicione uma lista ordenada com estes três itens de lista:

`flea treatment`
`thunder`
`other cats`

# --hints--

O elemento `ol` deve ter uma tag de abertura.As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert.exists(document.querySelector('ol'));
```

O elemento `ol` deve ter uma tag de fechamento.As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert.match(code, /<\/ol>/);
```

O elemento `ol` deve estar acima da tag de fechamento do segundo elemento `section`.Eles estão na ordem errada.

```js
assert.equal(document.querySelectorAll('main > section')[1]?.lastElementChild.nodeName, 'OL');
```

Os três elementos `li` devem estar dentro do elemento `ol`.

```js
assert.lengthOf(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ), 3
);
```

Você deve ter os três elementos `li` com os textos `flea treatment`, `thunder` e `other cats`, não importando a ordem.

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText?.trim().replace(/\s+/g, ' ').toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

Você deve ter apenas um elemento `ol`.

```js 
assert.lengthOf([...document.querySelectorAll('ol')], 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--
        <h3>Top 3 things cats hate:</h3>
        
--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```
