---
id: 5f07c98cdb9413cbd4b16750
title: Passo 19
challengeType: 0
dashedName: step-19
---

# --description--

É hora de adicionar uma nova seção. Adicione um segundo elemento `section` abaixo do elemento `section` existente.

# --hints--

O elemento `section` deve ter uma tag de abertura. As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert.isAtLeast(document.querySelectorAll('section').length, 2);
```

Você deve acrescentar apenas uma tag de abertura para `section`. Remova as tags adicionais.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
```

O elemento `section` deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert.isAtLeast(code.match(/<\/section>/g)?.length, 2);
```

Você deve acrescentar apenas uma tag de fechamento para `section`. Remova as tags adicionais.

```js
assert.lengthOf(code.match(/<\/section>/g), 2);
```

O segundo elemento `section` não deve estar dentro do primeiro elemento `section`.

```js
const el = document.querySelector('main > section');
assert.isNotNull(el);
const childrenOf1stSection = [
  ...el.children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.isEmpty(foundElems);
```

Os dois elementos `section` devem estar dentro das tags de abertura e fechamento do elemento `main`.

```js
const el = document.querySelector('main');
assert.isNotNull(el);
const childrenOfMain = [...el.children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.lengthOf(foundElems, 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
