---
id: 671b6e873249bb35c9debfcf
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

Adicione as palavras `See more` antes do elemento âncora, e  `in our gallery` após o elemento âncora.

# --hints--

You should use the existing anchor element, do not create a new one.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

O texto dentro do elemento de âncora, não pode ser mudado, ele precisa ser `cat photos`

```js
assert.strictEqual(document.querySelector('a')?.innerText?.replace(/\s+/g, ' ').trim(), "cat photos");
```

You should have `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery` in your code.

```js
assert.match(code, /See more\s+<a href="https:\/\/freecatphotoapp\.com">cat photos<\/a>\s+in our gallery/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
