---
id: 660eebd83100d37862268781
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

La tua variabile `character` attualmente non ha un valore. È possibile assegnare un valore utilizzando l'operatore <dfn>assegnazione</dfn> `=`. Per esempio:

```js
let hello = "Hello";
```

L'assegnazione di un valore a una variabile al momento della sua dichiarazione è nota come <dfn>inizializzazione</dfn>.

Inizializza la tua variabile `character` assegnandogli il valore `"Hello"` durante la sua dichiarazione.

# --hints--

Dovresti utilizzare l'operatore di assegnazione `=`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character\s*=/);
```

Dovresti usare la stringa `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /('|")Hello\1/);
```

Dovresti usare le doppie virgolette per la stringa `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

Dovresti assegnare `"Hello"` alla variabile `character`.

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character;
--fcc-editable-region--
```
