---
id: 691f7bbf76229bb97c60827d
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Nesta oficina, você vai praticar o uso de loops `for...of` construindo uma função que conta quantas vezes uma string aparece em um array de strings.

Antes de escrever essa função, você criará uma mais simples que registra cada caractere de uma string no console.

Comece definindo uma função vazia chamada `printCharacters` com o parâmetro `str`.

# --hints--

Você deve criar uma função `printCharacters`.

```js
assert.isFunction(printCharacters);
```

Sua função `printCharacters` deve ter um parâmetro chamado `str`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
