---
id: 6920c9cda1c7f056f75e4b1a
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Agora é hora de escrever seu primeiro loop `for...of`. Nas lições anteriores, você aprendeu que pode usar um loop `for...of` para iterar sobre valores de um iterável (por exemplo, strings ou arrays).

Aqui está um exemplo de um loop `for...of`:

```js
for (const num of [1, 2, 3]) {
  // code block to be executed
}
```

Adicione um loop `for...of` com um bloco de código vazio dentro da sua função. Ele deve iterar sobre cada caractere do argumento `str`.

# --hints--

Sua função `printCharacters` deve ter um loop `for...of` dentro do seu bloco de código.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functionprintCharacters\(str\)\{for\(\S+of\S+\)\{\S*?\}\}/);
```

Seu loop `for...of` deve iterar sobre o argumento `str` da sua função.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function printCharacters\(str\) \{\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+str/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
