---
id: 6920d189e36b6b60a6330fee
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Agora, dentro do loop, registre a variável `char` no console.

# --hints--

Você deve ter uma declaração `console.log` com `char` como argumento dentro do loop.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);
const loopCodeBlock = normalizedFunctionCode.match(/for\(constcharofstr\)\{(\S*?)}/)[1]

assert.match(loopCodeBlock, /console\.log\(char\)/);
```

Quando a função `printCharacters` é chamada com o argumento `'test'`, a declaração `console.log()` dentro da função deve ser chamada quatro vezes com os seguintes argumentos, na ordem: `'t'`, `'e'`, `'s'` e `'t'`. Em vez disso, foi chamada com os seguintes argumentos: `--fcc-actual--`.

```js
const consoleLogSpy = __helpers.spyOn(console, 'log');
printCharacters('test');
assert.deepEqual(consoleLogSpy.calls, [['t'], ['e'], ['s'], ['t']]);
consoleLogSpy.restore();
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
