---
id: 6921d90b5ce50185505d2004
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Como a função `getMatchedWordCount` precisa retornar uma contagem numérica, você vai começar declarando uma variável contador.

Crie uma variável chamada `count` usando a palavra-chave `let`, inicialize-a com o valor `0` e adicione uma declaração `return` que retorne a variável `count`.

# --hints--

Você deve declarar a variável `count` dentro da função `getMatchedWordCount`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /(const|let|var)count/);
```

Para permitir que a variável `count` seja incrementada depois, declare-a com a palavra-chave `let`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /letcount/);
```

Sua função deve conter uma atribuição do valor `0` para a variável `count`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /count=0/);
```

Sua função deve conter uma declaração `return` que retorne a variável `count`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /returncount/);
```

Sua função `getMatchedWordCount` deve retornar o valor `0`.

```js
assert.strictEqual(getMatchedWordCount(["foo", "bar"], "foo"), 0);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
