---
id: 6921d688b8d58581d34459b9
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Nos próximos passos, você vai criar uma função que conta quantas vezes uma string aparece em um array de strings.

Para começar, defina uma função vazia chamada `getMatchedWordCount` com os parâmetros `sentence` e `match` nessa ordem.

# --hints--

Você deve criar uma função `getMatchedWordCount`.

```js
assert.isFunction(getMatchedWordCount);
```

Sua função `getMatchedWordCount` deve ter dois parâmetros: `sentence` e `match`, nessa ordem.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
