---
id: 6921d688b8d58581d34459b9
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Nos próximos passos, você vai criar uma função que conta quantas vezes uma string aparece em um array de strings.

Para começar, defina uma função vazia chamada `getMatchedWordCount` com os parâmetros `sentence` e `match` nessa ordem.

# --hints--

Você deve criar uma função `getMatchedWordCount`.

```js
assert.isFunction(getMatchedWordCount);
```

Sua função `getMatchedWordCount` deve ter dois parâmetros: `sentence` e `match`, nessa ordem.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.getMatchedWordCount;
assert.lengthOf(parameters, 2);
assert.equal(parameters[0].toString(), 'sentence');
assert.equal(parameters[1].toString(), 'match');
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
