---
id: 6921dae10700e300f26994e3
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Para ver a saída da sua função `getMatchedWordCount` no estado atual, adicione uma declaração `console.log` logo abaixo dela. Dentro do `console.log()`, faça a chamada da função com o argumento `["I", "really", "really", "really", "like", "to", "code"]` para o parâmetro `sentence` e `"really"` para o parâmetro `match`.

# --hints--

Você deve fazer a chamada da função `getMatchedWordCount` dentro de uma declaração `console.log`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
assert.match(normalizedCode, /console\.log\(getMatchedWordCount\(\S*?\)\)/);
```

Você deve fazer a chamada da função `getMatchedWordCount` com o argumento `["I", "really", "really", "really", "like", "to", "code"]` para o parâmetro `sentence` e `"really"` para o parâmetro `match`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/)[0];
assert.match(normalizedConsoleLogCode, /\[('|"|`)I\1,('|"|`)really\2,('|"|`)really\3,('|"|`)really\4,('|"|`)like\5,('|"|`)to\6,('|"|`)code\7\],('|"|`)really\8/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
