---
id: 6921de8d2a111d0341cff702
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

Ótimo, sua função retorna um `0`, mas ainda não conta nada.

Para corrigir isso, dentro do `getMatchedWordCount`, crie um loop `for...of` com um bloco de código vazio que itere sobre cada palavra em `sentence`.

# --hints--

Sua função `getMatchedWordCount` deve ter um loop `for...of` dentro do seu bloco de código.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functiongetMatchedWordCount\(sentence,match\)\{letcount=0\;for\(\S+of\S+\)\{\S*?\}returncount;\}/);
```

Seu loop `for...of` deve iterar sobre o argumento `sentence` da sua função.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function getMatchedWordCount\(sentence, match\) \{\s*let count = 0\;\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*return count\;\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+sentence/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
--fcc-editable-region--
  
--fcc-editable-region--
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
