---
id: 66a9c14e3b34c719e34bae20
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Um Schülerdaten in die Tabelle einzufügen, müssen Sie die Tabellenzeilen- und Tabellendaten-Elemente verwenden.

Das <dfn>table data</dfn>-Element, `td`, wird verwendet, um eine Zelle in der Tabelle zu erstellen.

Hier ist ein Beispiel für die Verwendung des `td`-Elements für eine Tabelle mit Sportspielern:

```html
<tr>
  <td>1</td>
  <td>John Doe</td>
  <td>USA</td>
</tr>
```

Fügen Sie innerhalb Ihres Tabellenkörper-Elements ein Tabellenzeilen-Element, `tr`, mit drei Tabellendaten-Elementen, `td`, hinzu.

Das erste Tabellendaten-Element sollte den Nachnamen von `Davis` enthalten.

Das zweite Tabellendaten-Element sollte den Vornamen von `Alex` enthalten.

Das dritte Tabellendaten-Element sollte die Note von `54` enthalten.

# --hints--

Sie sollten ein `tr`-Element innerhalb Ihres `tbody`-Elements haben.

```js
assert.isNotNull(document.querySelector('tbody tr'));
```

Sie sollten ein `td`-Element mit dem Text `Davis` innerhalb des `tr`-Elements haben.

```js
assert.strictEqual(document.querySelector('tbody tr td')?.textContent.trim(), 'Davis');
```

Sie sollten ein `td`-Element mit dem Text `Alex` innerhalb des `tr`-Elements haben.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[1]?.textContent.trim(), 'Alex');
```

Sie sollten ein `td`-Element mit dem Text `54` innerhalb des `tr`-Elements haben.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[2]?.textContent.trim(), '54');
```

Sie sollten alle drei `td`-Elemente innerhalb des `tr`-Elements haben.

```js
assert.match(code, /<tr>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<\/tr>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
