---
id: 66a9ca099e3ecb1be10a2696
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Jetzt ist es Zeit, zwei weitere Schüler zur Tabelle hinzuzufügen.

Folgen Sie dem gleichen Muster wie in den vorherigen Schritten und fügen Sie eine vierte Tabellenzeile für einen Schüler hinzu. Verwenden Sie die folgenden Daten für die Tabellendaten-Elemente:

- Nachname: `Thompson`
- Vorname: `Jane`
- Note: `77`

Für die fünfte Tabellenzeile eines Schülers verwenden Sie die folgenden Daten für die Tabellendaten-Elemente:

- Nachname: `Williams`
- Vorname: `Natalie`
- Note: `83`

# --hints--

Sie sollten eine vierte `tr` innerhalb Ihres `tbody`-Elements haben.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(4)'));
```

Sie sollten ein `td`-Element mit dem Text `Thompson` innerhalb des vierten `tr`-Elements haben.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:first-child')?.textContent.trim(), 'Thompson');
```

Sie sollten ein `td`-Element mit dem Text `Jane` innerhalb des vierten `tr`-Elements haben.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(2)')?.textContent.trim(), 'Jane');
```

Sie sollten ein `td`-Element mit dem Text `77` innerhalb des vierten `tr`-Elements haben.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(3)')?.textContent.trim(), '77');
```

Sie sollten eine fünfte `tr` innerhalb Ihres `tbody`-Elements haben.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 5);
```

Sie sollten ein `td`-Element mit dem Text `Williams` innerhalb des fünften `tr`-Elements haben.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:first-child')?.textContent.trim(), 'Williams');
```

Sie sollten ein `td`-Element mit dem Text `Natalie` innerhalb des fünften `tr`-Elements haben.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(2)')?.textContent.trim(), 'Natalie');
```

Sie sollten ein `td`-Element mit dem Text `83` innerhalb des fünften `tr`-Elements haben.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(3)')?.textContent.trim(), '83');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
