---
id: 66a9bcf00f13a418368a272e
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Das Tabellenkopf-Element besteht aus einem <dfn>Tabellenzeilen</dfn>-Element, `tr`, das die <dfn>Tabellenkopf-Zellen</dfn>-Elemente `th` enthält.

Hier ist ein Beispiel, das die `tr`- und `th`-Elemente für eine Sporttabelle verwendet:

```html
<table>
  <caption>Football Scores</caption>
  <thead>
    <tr>
      <th>Team</th>
      <th>Wins</th>
      <th>Losses</th>
    </tr>
  </thead>
</table>
```

Fügen Sie innerhalb Ihres `thead`-Elements ein `tr`-Element hinzu.

Fügen Sie innerhalb Ihres `tr`-Elements drei `th`-Elemente hinzu.

Das erste `th`-Element sollte den Text `Last Name` enthalten. Das zweite `th`-Element sollte den Text `First Name` enthalten. Das dritte `th`-Element sollte den Text `Grade` enthalten.

# --hints--

Sie sollten ein öffnendes `tr`-Tag haben.

```js
assert.match(code, /<tr>/i);
```

Sie sollten ein schließendes `tr`-Tag haben.

```js
assert.match(code, /<\/tr>/i);
```

Ihr `tr`-Element sollte sich innerhalb Ihres `thead`-Elements befinden.

```js
const trEl = document.querySelector('tr');
assert.strictEqual(trEl?.parentElement?.tagName, 'THEAD');
```

Ihr erstes `th`-Element sollte den Text von `Last Name` haben.

```js
const thEl = document.querySelectorAll('th')[0];
assert.strictEqual(thEl?.textContent.trim(), 'Last Name');
```

Ihr zweites `th`-Element sollte den Text von `First Name` haben.

```js
const thEl = document.querySelectorAll('th')[1];
assert.strictEqual(thEl?.textContent.trim(), 'First Name');
```

Ihr drittes `th`-Element sollte den Text von `Grade` haben.

```js
const thEl = document.querySelectorAll('th')[2];
assert.strictEqual(thEl?.textContent.trim(), 'Grade');
```

Alle Ihre `th`-Elemente sollten sich innerhalb Ihres `tr`-Elements befinden.

```js
assert.match(code, /<tr>.*<th>.*<\/th>.*<th>.*<\/th>.*<th>.*<\/th>.*<\/tr>/is);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
      --fcc-editable-region--
        
      --fcc-editable-region--
      </thead>
    </table>
  </body>
</html>
```
