---
id: 66a9bb8578459a175432e7d0
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Für den ersten Abschnitt der Tabelle möchten Sie den Kopfzeileninhalt gruppieren, der die Spaltenbezeichnungen für den Vornamen, Nachnamen und die Note der Abschlussprüfung des Schülers darstellt.

Das <dfn>table head</dfn>-Element, `thead`, wird verwendet, um den Kopfzeileninhalt in einer Tabelle zu gruppieren.

Hier ist ein Beispiel mit dem `thead`-Element:

```html
<table>
  <thead>
    <!-- header content goes here -->
  </thead>
</table>
```

Fügen Sie unter Ihrem `caption`-Element ein table head-Element hinzu.

# --hints--

Sie sollten ein öffnendes `thead`-Tag haben.

```js
assert.match(code, /<thead>/i);
```

Sie sollten ein schließendes `thead`-Tag haben.

```js
assert.match(code, /<\/thead>/i);
```

Ihr `thead`-Element sollte sich innerhalb Ihres `table`-Elements befinden.

```js
const theadEl = document.querySelector('thead');
assert.strictEqual(theadEl?.parentElement?.tagName, 'TABLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
