---
id: 66a9cb26ec6bd41cf6c82bc5
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Der letzte Abschnitt, der zur Tabelle hinzugefügt werden soll, ist das <dfn>table foot</dfn>-Element, `tfoot`. Das table foot-Element wird verwendet, um die Durchschnittsnote aller Schüler in der Tabelle anzuzeigen.

Fügen Sie der Tabelle ein `tfoot`-Element hinzu.

Fügen Sie innerhalb des `tfoot`-Elements ein `tr`-Element hinzu.

Fügen Sie innerhalb des `tr`-Elements zwei `td`-Elemente hinzu.

Das erste `td`-Element sollte den Text `Average Grade` enthalten. Das zweite `td`-Element sollte den Text `78.8` enthalten.

# --hints--

Sie sollten ein öffnendes `tfoot`-Tag haben.

```js
assert.match(code, /<tfoot>/);
```

Sie sollten ein schließendes `tfoot`-Tag haben.

```js
assert.match(code, /<\/tfoot>/);
```

Sie sollten ein `tr`-Element innerhalb Ihres `tfoot`-Elements haben.

```js
assert.isNotNull(document.querySelector('tfoot tr'));
```

Ihr erstes `td`-Element sollte den Text `Average Grade` enthalten. Stellen Sie sicher, dass es sich innerhalb des `tr`-Elements befindet.

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child')?.textContent.trim(), 'Average Grade');
```

Ihr zweites `td`-Element sollte den Text `78.8` enthalten. Stellen Sie sicher, dass es sich innerhalb des `tr`-Elements befindet.

```js
assert.strictEqual(document.querySelector('tfoot tr td:nth-of-type(2)')?.textContent.trim(), '78.8');
```

Ihr `tr`-Element sollte zwei `td`-Elemente enthalten.

```js
assert.strictEqual(document.querySelectorAll('tfoot tr td').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
