---
id: 66a9cd4ffa48cf1ef7333640
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Ihre Tabelle ist fast fertig, aber es gibt noch eine letzte Sache hinzuzufügen.

Es wäre schön, wenn das `td`-Element, das für das `Average Grade` verwendet wird, sich über zwei Spalten erstrecken würde, statt nur über eine.

Dafür können Sie das <dfn>colspan</dfn>-Attribut am `td`-Element verwenden.

Das `colspan`-Attribut gibt an, über wie viele Spalten sich eine Zelle erstrecken soll.

Hier ist ein Beispiel für die Verwendung des `colspan`-Attributs in einer Sporttabelle:

```html
<tr>
  <td colspan="3">Total Points</td>
</tr>
```

Fügen Sie im öffnenden `td`-Tag das `colspan`-Attribut hinzu und setzen Sie es auf `"2"`.

Und mit dieser Änderung ist Ihre Tabelle komplett!

# --hints--

Sie sollten den Inhalt des bestehenden `td`-Elements nicht verändern.

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child').textContent.trim(), 'Average Grade');
```

Ihr `td`-Element sollte ein `colspan`-Attribut haben.

```js
assert.isTrue(document.querySelector('tfoot tr td:first-child').hasAttribute('colspan'));
```

Ihr `td`-Element sollte ein `colspan`-Attribut mit dem Wert `"2"` haben.

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child').getAttribute('colspan'), '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>

      <tfoot>
        <tr>
          --fcc-editable-region--
          <td>Average Grade</td>
          --fcc-editable-region--
          <td>78.8</td>
        </tr>
      </tfoot>
    </table>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>

      <tfoot>
        <tr>
          <td colspan="2">Average Grade</td>
          <td>78.8</td>
        </tr>
      </tfoot>
    </table>
  </body>
</html>
```
