---
id: 66a9b48a211a73155621d0df
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop üben Sie den Umgang mit HTML-Tabellen, indem Sie eine Prüfungstabelle für eine Gruppe von Studierenden erstellen.

In vorherigen Lektionen haben Sie gelernt, wie Sie mit dem `table`-Element tabellarische Daten darstellen.

Verschachteln Sie innerhalb Ihres `body`-Elements ein `table`-Element.

# --hints--

Sie sollten ein öffnendes `table`-Tag haben.

```js
assert.match(code, /<table>/i);
```

Sie sollten ein schließendes `table`-Tag haben.

```js
assert.match(code, /<\/table>/i);
```

Ihr `table`-Element sollte sich innerhalb Ihres `body`-Elements befinden.

```js
const tableEl = document.querySelector('table');
assert.strictEqual(tableEl?.parentElement?.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
