---
id: 66a9c7a4521d2b1b1ec6dcf0
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Jetzt ist es Zeit, zwei weitere Studierende zur Tabelle hinzuzufügen.

Folgen Sie dem gleichen Muster wie im vorherigen Schritt und fügen Sie eine zweite Tabellenzeile für einen Studierenden hinzu. Verwenden Sie die folgenden Daten für die Tabellendaten-Elemente:

- Nachname: `Doe`
- Vorname: `Samantha`
- Note: `92`

Für die dritte Tabellenzeile eines Studierenden verwenden Sie die folgenden Daten für die Tabellendaten-Elemente:

- Nachname: `Rodriguez`
- Vorname: `Marcus`
- Note: `88`

# --hints--

Sie sollten eine zweite `tr` innerhalb Ihres `tbody`-Elements haben.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(2)'));
```

Sie sollten ein `td`-Element mit dem Text `Doe` innerhalb des zweiten `tr`-Elements haben.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:first-child')?.textContent.trim(), 'Doe');
```

Sie sollten ein `td`-Element mit dem Text `Samantha` innerhalb des zweiten `tr`-Elements haben.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:nth-child(2)')?.textContent.trim(), 'Samantha');
```

Sie sollten ein `td`-Element mit dem Text `92` innerhalb des zweiten `tr`-Elements haben.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:last-child')?.textContent.trim(), '92');
```

Sie sollten eine dritte `tr` innerhalb Ihres `tbody`-Elements haben.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 3);
```

Sie sollten ein `td`-Element mit dem Text `Rodriguez` innerhalb des dritten `tr`-Elements haben.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:first-child')?.textContent.trim(), 'Rodriguez');
```

Sie sollten ein `td`-Element mit dem Text `Marcus` innerhalb des dritten `tr`-Elements haben.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:nth-child(2)')?.textContent.trim(), 'Marcus');
```

Sie sollten ein `td`-Element mit dem Text `88` innerhalb des dritten `tr`-Elements haben.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:last-child')?.textContent.trim(), '88');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
