---
id: 66a9b8f14b963916a3baa732
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Um eine Bildunterschrift zu einem `table` hinzuzufügen, können Sie das <dfn>table caption</dfn>-Element verwenden.

Hier ist ein Beispiel mit dem `caption`-Element:

```html
<table>
  <caption>Football Scores</caption>
</table>
```

Verschachteln Sie in Ihrem `table`-Element ein `caption`-Element mit dem Text `Calculus Final Exam Grades`.

# --hints--

Sie sollten ein öffnendes `caption`-Tag haben.

```js
assert.match(code, /<caption>/i);
```

Sie sollten ein schließendes `caption`-Tag haben.

```js
assert.match(code, /<\/caption>/i);
```

Ihr `caption`-Element sollte sich innerhalb Ihres `table`-Elements befinden.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.parentElement.tagName, 'TABLE');
```

Ihr `caption` sollte den Text `Calculus Final Exam Grades` enthalten.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.innerText.trim(), 'Calculus Final Exam Grades');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
    --fcc-editable-region--
      
    --fcc-editable-region--      
    </table>
  </body>
</html>
```
