---
id: 66a9c006e4f58819396605b3
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Nachdem Sie den Kopfbereich abgeschlossen haben, ist es Zeit, den <dfn>Tabellenkörper</dfn>, `tbody`, hinzuzufügen. Der Tabellenkörper wird alle Studentennamen und ihre Noten darstellen.

Fügen Sie Ihrer Tabelle ein Tabellenkörper-Element hinzu.

# --hints--

Sie sollten ein öffnendes `tbody`-Tag haben.

```js
assert.match(code, /<tbody>/i);
```

Sie sollten ein schließendes `tbody`-Tag haben.

```js
assert.match(code, /<\/tbody>/i);
```

Ihr öffnendes `tbody`-Tag sollte vor dem schließenden `table`-Tag stehen.

```js
assert.match(code, /<tbody>[\s\S]*<\/tbody>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
