---
id: 61437d575fb98f57fa8f7f36
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

標準的なHTMLボイラープレートから始めてください。`DOCTYPE` 宣言、ページが英語であることを指定する`html` 要素、`head` 要素、そして`body` 要素を追加してください。

`<meta>` 要素内に適切な`charset` を持つ`<meta>` タグと、モバイルレスポンシブ対応のための`head` タグを追加してください。

# --hints--

コードには`DOCTYPE` 参照が含まれている必要があります。

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 参照の後にスペースを入れてください。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

ドキュメントタイプは`html` に設定してください。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

型の後に`DOCTYPE` で`>` 宣言を閉じてください。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

開く`<html>` タグがあり、それには`lang` の`en` が必要です。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 要素には終了タグが必要です。

```js
assert(code.match(/<\/html\s*>/));
```

`DOCTYPE` 宣言はHTMLの最初に置いてください。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

開く`<head>` タグが必要です。

```js
assert(code.match(/<head\s*>/i));
```

閉じる`</head>` タグが必要です。

```js
assert(code.match(/<\/head\s*>/i));
```

開く`<body>` タグが必要です。

```js
assert(code.match(/<body\s*>/i));
```

閉じる`</body>` タグが必要です。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` と`body` 要素は兄弟要素である必要があります。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 要素は`html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 要素は`html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

`meta` 要素が2つ必要です。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

1つの`meta` 要素には`name` が`viewport` に、`content` が`width=device-width, initial-scale=1.0` に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

もう1つの`meta` 要素には`charset` 属性が`UTF-8` に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
