---
id: 614385513d91ae5c251c2052
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

テキスト `title` を持つ `Magazine` 要素、`link` フォントのスタイルシート用の `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` 要素、`link` FontAwesome スタイルシート用の `https://use.fontawesome.com/releases/v5.8.2/css/all.css`、そして `link` スタイルシート用の `./styles.css` を追加してください。

フォントのスタイルシートは `Anton`、`Baskervville`、`Raleway` の3つの異なるフォントを読み込みます。

# --hints--

コードには3つの `link` 要素が必要です。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

`link` 要素は `head` 要素の中に入っている必要があります。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

3つの `link` 要素には `rel` 属性があり、その値は `stylesheet` でなければなりません。

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

1つのリンク要素には `href` が `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` に設定されている必要があります。

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

1つのリンク要素には `href` が `https://use.fontawesome.com/releases/v5.8.2/css/all.css` に設定されている必要があります。

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

1つの `link` 要素には `href` が `styles.css` に設定されている必要があります。

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

コードには `title` 要素が必要です。

```js
const title = document.querySelector('title');
assert.isNotNull(title);
```

コードには1つの `title` 要素が必要です。

```js
const title = document.querySelectorAll('title');
assert.strictEqual(title?.length, 1);
```

`title` 要素は `head` 要素の中に入っている必要があります。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('title'));
```

プロジェクトのタイトルは `Magazine` にしてください。タイトルの大文字・小文字やスペルは正確に一致させる必要があります。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Magazine');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
