---
id: 614e0e588f0e8a772a8a81a6
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

現在の画像はスペースを多く占めています。作業中の内容をよりよく見えるように、`width` 要素に `img` 属性を追加し、その値を `400` にしてください。

CSSの作業が終わったら、この属性は後で削除します。

# --hints--

`img` 要素には `width` 属性を `400` に設定してください。

```js
assert(document.querySelector('img')?.getAttribute('width') === '400');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
