---
id: 67ed03ac474c48692f41749e
title: Costruire una tabella hash
challengeType: 27
dashedName: build-a-hash-table
saveSubmissionToDB: true
---

# --description--

In questo laboratorio costruirai una tabella hash da zero. Una tabella hash è una struttura dati che archivia coppie chiave-valore. Una tabella hash funziona prendendo la chiave come input e poi eseguendo l'hash di questa chiave secondo una specifica funzione di hashing.

Per questo laboratorio, la funzione di hashing sarà semplice: sommerà i valori Unicode di ogni carattere nella chiave. Il valore hash sarà poi usato come chiave effettiva per archiviare il valore associato. Lo stesso valore hash sarà usato anche per recuperare ed eliminare il valore associato alla chiave.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti definire una classe chiamata `HashTable` con un attributo `collection` inizializzato a un dizionario vuoto quando viene creata una nuova istanza di `HashTable`. Il dizionario `collection` dovrebbe archiviare coppie chiave-valore basate sul valore hash della chiave.

2. La classe `HashTable` dovrebbe avere quattro metodi di istanza: `hash`, `add`, `remove` e `lookup`.

3. Il metodo `hash` dovrebbe:

    - Prendere una stringa come parametro.
    - Restituire un valore hash calcolato come la somma dei valori Unicode (ASCII) di ogni carattere nella stringa. Puoi usare la funzione `ord` per questo calcolo.

4. Il metodo `add` dovrebbe:
    - Prendere due argomenti che rappresentano una coppia chiave-valore e calcolare l'hash della chiave.
    - Usare il valore hash calcolato come chiave per archiviare un dizionario contenente la coppia chiave-valore dentro il dizionario `collection`.
    - Se più chiavi producono lo stesso valore hash, le loro coppie chiave-valore dovrebbero essere archiviate nel dizionario annidato esistente sotto lo stesso valore hash.

5. Il metodo `remove` dovrebbe:

    - Prendere una chiave come argomento e calcolarne l'hash.
    - Verificare se la chiave esiste nella collezione.
    - Rimuovere la coppia chiave-valore corrispondente dalla tabella hash.
    - Se la chiave non esiste nella collezione, non dovrebbe generare errori né rimuovere nulla.

6. Il metodo `lookup` dovrebbe:

    - Prendere una chiave come argomento.
    - Calcolare l'hash della chiave e restituire il valore corrispondente archiviato nella tabella hash.
    - Se la chiave non esiste nella collezione, dovrebbe restituire `None`.

# --hints--

Dovresti definire una classe `HashTable`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_class("HashTable")
    `)) 
})
```

Quando viene creata una nuova istanza della classe `HashTable`, il suo attributo `collection` dovrebbe essere inizializzato a un dizionario vuoto.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        assert ht.collection == {}
    `) 
})
```

La classe `HashTable` dovrebbe avere un metodo `hash`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('hash')`))}})
```

Il metodo `hash` dovrebbe prendere una stringa come parametro.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.hash)
    len(sig.parameters) == 1
  `))
})
```

Il metodo `hash` dovrebbe prendere una stringa come argomento e restituire la somma dei valori Unicode di ogni carattere nella stringa.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        hash_result = ht.hash("fcc")
        assert hash_result == 300

        # prevent hardcoding
        assert ht.hash("golf") == 424
        assert ht.hash("read") == 412
`) })
```

La classe `HashTable` dovrebbe avere un metodo `add`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('add')`))}})
```

Il metodo `add` dovrebbe prendere una chiave e un valore come parametri.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.add)
    len(sig.parameters) == 2
  `))
})
```

La classe `HashTable` dovrebbe avere un metodo `remove`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('remove')`))}})
```

Il metodo `remove` dovrebbe prendere una chiave come parametro.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.remove)
    len(sig.parameters) == 1
  `))
})
```

Quando provi a rimuovere una chiave che non esiste nella collezione, non dovrebbe generare errori né rimuovere nulla.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        ht.add("rose", "flower")
        index = ht.hash("rose")
        original = ht.collection.copy()

        ht.remove("tulip")
        ht.remove("sore")

        assert ht.collection == original
        assert "rose" in ht.collection[index]
        assert "tulip" not in ht.collection.get(index, {})
    `) 
})
```

Se più chiavi hanno lo stesso valore hash, il metodo `remove` dovrebbe eliminare solo la specifica coppia chiave-valore e non l'intero dizionario a quell'indice.

```js
({
  test: () => runPython(`
    ht = HashTable()
    ht.add("rose", "flower")
    ht.add("sore", "pain")  # "rose" and "sore" both hash to the same index

    index = ht.hash("rose")
    ht.remove("rose")

    assert index in ht.collection
    assert "rose" not in ht.collection[index]
    assert "sore" in ht.collection[index]
    assert ht.collection[index]["sore"] == "pain"
  `)
})
```

La classe `HashTable` dovrebbe avere un metodo `lookup`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('lookup')`))}})
```

Il metodo `lookup` dovrebbe prendere una chiave come parametro.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.lookup)
    len(sig.parameters) == 1
  `))
})
```

`HashTable().hash('golf')` dovrebbe restituire `424`.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        hash_result = ht.hash("golf")
        assert hash_result == 424

        # prevent hardcoding
        assert ht.hash("dear") == 412
        assert ht.hash("cat") == 312
    `) 
})
```

`HashTable().add('golf', 'sport')` dovrebbe aggiungere la coppia chiave-valore alla collezione alla chiave `424`.

```js

({ 
    test: () => runPython(`
    ht = HashTable()
    ht.add("golf", "sport")
    expected_value = {424: {'golf': 'sport'}}

    assert ht.collection == expected_value
    `) 
})
```

`HashTable().add('dear', 'friend')` e `HashTable().add('read', 'book')` dovrebbero aggiungere entrambe le coppie chiave-valore alla collezione all'indice `412` come dizionario annidato.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("dear", "friend")
    ht.add("read", "book")

    expected_value = {
        "dear": "friend",
        "read": "book"
    }

    assert ht.collection.get(412) == expected_value
    `) 
})
```

Quando una chiave esiste nella tabella hash, il metodo `remove()` dovrebbe rimuovere la chiave data e il valore corrispondente dalla collezione.

```js
({ 
    test: () => runPython(`
    ht = HashTable()
    ht.add("golf", "sport")

    expected_value_before_removal = {
        "golf": "sport"
    }

    index = ht.hash("golf")

    assert ht.collection.get(index) == expected_value_before_removal

    ht.remove("golf")

    assert "golf" not in ht.collection.get(index, {})
    `) 
})
```

Quando la coppia chiave-valore `'golf', 'sport'` esiste nella tabella hash, `HashTable().lookup('golf')` dovrebbe restituire `sport`.

```js
({ 
    test: () => runPython(`
    ht = HashTable()

    ht.add("golf", "sport")

    expected_value = "sport"
    assert ht.lookup("golf") == expected_value
    `) 
})

```

Quando la coppia chiave-valore `'golf', 'sport'` non esiste nella collezione, `HashTable().lookup('golf')` dovrebbe restituire `None`.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        assert ht.lookup("golf") is None
`) })

```

Quando la chiave `'fcc'` esiste nella collezione, `HashTable().lookup('cfc')` dovrebbe restituire `None`.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        ht.add("fcc", "coding")
        assert ht.lookup("cfc") is None
`) })
```

Quando aggiungi `('rose', 'flower')` alla tabella hash, il suo attributo `collection` dovrebbe apparire così: `{ 441: { 'rose': 'flower' }}`.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("rose", "flower")

    expected_value = {
        "rose": "flower"
    }

    assert ht.collection.get(441) == expected_value

    # prevent hardcoding

    ht2 = HashTable()

    ht2.add("kebab", "food")
    expected_value = {
        "kebab": "food"
    }

    assert ht2.collection.get(501) == expected_value
    `)
})
```

Quando aggiungi una chiave che ha lo stesso valore hash di una chiave esistente, come `fcc` e `cfc`, `collection` dovrebbe apparire così: `{ 300: { 'fcc': 'coding', 'cfc':  'chemical' }}`.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("fcc", "coding")
    ht.add("cfc", "chemical")

    expected_value = {
        "fcc": "coding",
        "cfc": "chemical"
    }

    assert ht.collection.get(300) == expected_value

    # prevent hardcoding

    ht2 = HashTable()
    ht2.add("cat", "animal")
    ht2.add("act", "verb")

    expected_value = {
        "cat": "animal",
        "act": "verb"
    }

    assert ht2.collection.get(312) == expected_value
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class HashTable:
    def __init__(self):
        self.collection = {}

    def hash(self, string):
        hashed = 0
        for char in string:
            hashed += ord(char)  
        return hashed

    def add(self, key, val):
        the_hash = self.hash(key)
        if the_hash not in self.collection:
            self.collection[the_hash] = {}
        self.collection[the_hash][key] = val

    def remove(self, key):
        the_hash = self.hash(key)
        if the_hash in self.collection and key in self.collection[the_hash]:
            del self.collection[the_hash][key]
            if not self.collection[the_hash]: 
                del self.collection[the_hash]

    def lookup(self, key):
        the_hash = self.hash(key)
        if the_hash in self.collection and key in self.collection[the_hash]:
            return self.collection[the_hash][key]
        return None
```
