---
id: 587d824a367417b2b2512c45
title: Anonymes Messageboard
challengeType: 4
forumTopicId: 301568
dashedName: anonymous-message-board
---

# --description--

Erstellen Sie eine Full-Stack-JavaScript-App, die funktional ähnlich ist wie diese: <a href="https://anonymous-message-board.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://anonymous-message-board.freecodecamp.rocks/</a>.

Bei der Arbeit an diesem Projekt schreiben Sie Ihren Code mit einer der folgenden Methoden:

-   Klonen Sie <a href="https://github.com/freeCodeCamp/boilerplate-project-messageboard/" target="_blank" rel="noopener noreferrer nofollow">dieses GitHub-Repository</a> und vervollständigen Sie Ihr Projekt lokal.
-   Verwenden Sie einen Website-Builder Ihrer Wahl, um das Projekt abzuschließen. Stellen Sie sicher, dass Sie alle Dateien aus unserem GitHub-Repository einbinden.

# --instructions--

1.  Setzen Sie `NODE_ENV`, um ohne Anführungszeichen zu testen, wenn Sie bereit sind, Tests zu schreiben, und `.env` auf Ihre Datenbank-Verbindungszeichenfolge
2.  Es wird empfohlen, Controller/Handler zu erstellen und das Routing in `routes/api.js` zu verwalten
3.  Sie fügen alle Sicherheitsfunktionen in `server.js` hinzu

Schreiben Sie die folgenden Tests in `tests/2_functional-tests.js`:

-   Erstellen eines neuen Threads: POST-Anfrage an `/api/threads/{board}`
-   Anzeigen der 10 neuesten Threads mit jeweils 3 Antworten: GET-Anfrage an `/api/threads/{board}`
-   Löschen eines Threads mit falschem Passwort: DELETE-Anfrage an `/api/threads/{board}` mit ungültigem `delete_password`
-   Löschen eines Threads mit richtigem Passwort: DELETE-Anfrage an `/api/threads/{board}` mit gültigem `delete_password`
-   Melden eines Threads: PUT-Anfrage an `/api/threads/{board}`
-   Erstellen einer neuen Antwort: POST-Anfrage an `/api/replies/{board}`
-   Anzeigen eines einzelnen Threads mit allen Antworten: GET-Anfrage an `/api/replies/{board}`
-   Löschen einer Antwort mit falschem Passwort: DELETE-Anfrage an `/api/replies/{board}` mit ungültigem `delete_password`
-   Löschen einer Antwort mit richtigem Passwort: DELETE-Anfrage an `/api/replies/{board}` mit gültigem `delete_password`
-   Melden einer Antwort: PUT-Anfrage an `/api/replies/{board}`

# --hints--

Sie können Ihr eigenes Projekt bereitstellen, nicht die Beispiel-URL.

```js
  assert(
    !/.*\/anonymous-message-board\.freecodecamp\.rocks/.test(
      code
    )
  );
```

Erlauben Sie nur, dass Ihre Seite in einem iFrame auf Ihren eigenen Seiten geladen wird.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['x-frame-options']?.includes('SAMEORIGIN'));
```

Verhindern Sie DNS-Prefetching.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['x-dns-prefetch-control']?.includes('off'));
```

Erlauben Sie nur, dass Ihre Seite den Referrer für Ihre eigenen Seiten sendet.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['referrer-policy']?.includes('same-origin'));
```

Sie können eine POST-Anfrage an `/api/threads/{board}` mit Formulardaten einschließlich `text` und `delete_password` senden. Der gespeicherte Datenbankeintrag enthält mindestens die Felder `_id`, `text`, `created_on` (Datum & Uhrzeit), `bumped_on` (Datum & Uhrzeit, beginnt gleich wie `created_on`), `reported` (Boolean), `delete_password` und `replies` (Array).

```js
  const date = new Date();
  const text = `fcc_test_${date}`;
  const deletePassword = 'delete_me';
  const data = { text, delete_password: deletePassword };
  const url = code;
  const res = await fetch(url + '/api/threads/fcc_test', {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  if (res.ok) {
    const checkData = await fetch(url + '/api/threads/fcc_test');
    const parsed = await checkData.json();
    try {
      assert.equal(parsed[0].text, text);
      assert.isNotNull(parsed[0]._id);
      assert.equal(new Date(parsed[0].created_on).toDateString(), date.toDateString());
      assert.equal(parsed[0].bumped_on, parsed[0].created_on);
      assert.isArray(parsed[0].replies);
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Sie können eine POST-Anfrage an `/api/replies/{board}` mit Formulardaten einschließlich `text`, `delete_password` und `thread_id` senden. Dies aktualisiert das `bumped_on`-Datum auf das Datum des Kommentars. Im `replies`-Array des Threads wird ein Objekt mit mindestens den Eigenschaften `_id`, `text`, `created_on`, `delete_password` und `reported` gespeichert.

```js
  const url = code;
  const body = await fetch(url + '/api/threads/fcc_test');
  const thread = await body.json();

  const date = new Date();
  const text = `fcc_test_reply_${date}`;
  const delete_password = 'delete_me';
  const thread_id = thread[0]._id;
  const replyCount = thread[0].replies.length;

  const data = { text, delete_password, thread_id };
  const res = await fetch(url + '/api/replies/fcc_test', {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  if (res.ok) {
    const checkData = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);
    const parsed = await checkData.json();
    try {
      assert.equal(parsed.replies.length, replyCount + 1);
      assert.equal(parsed.replies[0].text, text);
      assert.equal(parsed._id, thread_id);
      assert.equal(parsed.bumped_on, parsed.replies[0].created_on);
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Sie können eine GET-Anfrage an `/api/threads/{board}` senden. Zurückgegeben wird ein Array der 10 neuesten Threads mit jeweils nur den 3 neuesten Antworten. Die Felder `reported` und `delete_password` werden nicht an den Client gesendet.

```js
  const url = code;
  const res = await fetch(url + '/api/threads/fcc_test');

  if (res.ok) {
    const threads = await res.json();
    try {
      assert.equal(res.status, 200);
      assert.isAtMost(threads.length, 10);
      for (let i = 0; i < threads.length; i++) {
        assert.containsAllKeys(threads[i], ["_id", "text", "created_on", "bumped_on", "replies"]);
        assert.isAtMost(threads[i].replies.length, 3);
        assert.notExists(threads[i].delete_password);
        assert.notExists(threads[i].reported);
        for (let j = 0; j < threads[i].replies.length; j++) {
          assert.notExists(threads[i].replies[j].delete_password);
          assert.notExists(threads[i].replies[j].reported);
        }
      }
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Sie können eine GET-Anfrage an `/api/replies/{board}?thread_id={thread_id}` senden. Zurückgegeben wird der gesamte Thread mit allen Antworten, ebenfalls ohne die gleichen Felder wie im vorherigen Test.

```js
  const url = code;
  let res = await fetch(url + '/api/threads/fcc_test');
  const threads = await res.json();
  const thread_id = threads[0]._id;
  res = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);

  if (res.ok) {
    const thread = await res.json();
    try {
      assert.equal(res.status, 200);
      assert.isObject(thread);
      assert.containsAllKeys(thread, ["_id", "text", "created_on", "bumped_on", "replies"]);
      assert.isArray(thread.replies);
      assert.notExists(thread.delete_password);
      assert.notExists(thread.reported);
      for (let i = 0; i < thread.replies.length; i++) {
        assert.notExists(thread.replies[i].delete_password);
        assert.notExists(thread.replies[i].reported);
      }
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Sie können eine DELETE-Anfrage an `/api/threads/{board}` senden und `thread_id` sowie `delete_password` übergeben, um den Thread zu löschen. Zurückgegeben wird der String `incorrect password` oder `success`.

```js
  const url = code;
  let res = await fetch(url + '/api/threads/fcc_test');
  const threads = await res.json();
  const thread_id = threads[0]._id;
  let data = { thread_id, delete_password: "wrong_password" };
  const res_invalid = await fetch(url + '/api/threads/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  data = { thread_id, delete_password: "delete_me" };
  res = await fetch(url + '/api/threads/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const deleted = await res.text();
    const not_deleted = await res_invalid.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(deleted, "success");
      assert.equal(not_deleted, "incorrect password");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Sie können eine DELETE-Anfrage an `/api/replies/{board}` senden und `thread_id`, `reply_id` sowie `delete_password` übergeben. Zurückgegeben wird der String `incorrect password` oder `success`. Bei Erfolg wird der Text von `reply_id` in `[deleted]` geändert.

```js
  const url = code;

  const thread_data = {
    text: "fcc_test_thread",
    delete_password: "delete_me",
  };
  await fetch(`${url}/api/threads/fcc_test`, {
    method: "POST",
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(thread_data)
  });
  let res = await fetch(`${url}/api/threads/fcc_test`);
  let threads = await res.json();
  const thread_id = threads[0]._id;
  
  const reply_data = { thread_id, text: "fcc_test_reply", delete_password: "delete_me" };
  await fetch(`${url}/api/replies/fcc_test`, {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(reply_data)
  });
  res = await fetch(`${url}/api/threads/fcc_test`);
  threads = await res.json();
  const reply_id = threads[0].replies[0]._id;

  const data = { thread_id, reply_id, delete_password: "delete_me" };
  res = await fetch(url + '/api/replies/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const deleted = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(deleted, "success");
      res = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);
      const thread = await res.json();
      assert.equal(thread._id, thread_id);
      assert.equal(thread.replies[0]._id, reply_id);
      assert.equal(thread.replies[0].text, "[deleted]");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Sie können eine PUT-Anfrage an `/api/threads/{board}` senden und `thread_id` übergeben. Zurückgegeben wird der String `reported`. Der `reported`-Wert von `thread_id` wird in `true` geändert.

```js
  const url = code;

  let res = await fetch(`${url}/api/threads/fcc_test`);
  const threads = await res.json();
  const thread_id = threads[0]._id;
  const data = { thread_id };
  
  res = await fetch(`${url}/api/threads/fcc_test`, {
    method: 'PUT',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const reported = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(reported, "reported");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Sie können eine PUT-Anfrage an `/api/replies/{board}` senden und `thread_id` sowie `reply_id` übergeben. Zurückgegeben wird der String `reported`. Der `reported`-Wert von `reply_id` wird in `true` geändert.

```js
  const url = code;

  let res = await fetch(`${url}/api/threads/fcc_test`);
  const threads = await res.json();
  const thread_id = threads[0]._id;
  const reply_id = threads[0].replies[0]._id;
  const data = { thread_id, reply_id };
  
  res = await fetch(`${url}/api/replies/fcc_test`, {
    method: 'PUT',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const reported = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(reported, "reported");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Alle 10 funktionalen Tests sind abgeschlossen und bestanden.

```js
  const tests = await fetch(code + '/_api/get-tests');
  const parsed = await tests.json();
  assert.isTrue(parsed.length >= 10);
  parsed.forEach((test) => {
    assert.equal(test.state, 'passed');
    assert.isAtLeast(test.assertions.length, 1);
  });
```

