---
id: 6894c1b715b8d18a18ef3203
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika warsha hii, utapata mazoezi na darasa la vitu na vitu vya Python kwa kujenga orodha ya vyombo vya muziki.

Ili kufafanua darasa la vitu katika Python, tumia neno kuu la `class` ikifuatiwa na jina la darasa na alama ya nukta mbili.

```python
class ClassName:
    # class contents go here
```

Kumbuka kuwa katika Python, majina ya darasa la vitu kwa kawaida hutumia PascalCase, ambayo inamaanisha kila neno huanza na herufi kubwa bila nafasi kati ya maneno.

Anza kwa kuunda darasa la vitu linaloitwa `MusicalInstrument`. Ongeza `pass` kama mwili wa darasa la vitu kwa sasa.

# --hints--

Unapaswa kuunda darasa la vitu linaloitwa `MusicalInstrument`.

```js
({
 test: () => assert(runPython(`_Node(_code).has_class("MusicalInstrument")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
