---
id: 6895e0bdb0348bf46952ab67
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Sasa ni wakati wa kuunda matukio ya darasa lako la `MusicalInstrument`.

Ili kuunda tukio la darasa, unaitisha jina la darasa kama kitendakazi, ukipitisha hoja zinazohitajika kwa njia ya `__init__` (isipokuwa `self`, ambayo inasimamiwa moja kwa moja na Python).

```python
my_instance = ClassName(arg1, arg2)
```

Unapounda tukio hili, Python huita njia yako ya `__init__` na kuhifadhi thamani kama sifa za tukio ambazo ni za tukio hili pekee.

Nje ya ufafanuzi wa darasa, unda tukio linaloitwa `instrument_1` kwa `Oboe` ambayo ni chombo cha `woodwind`.

# --hints--

Unapaswa kuunda tukio la darasa la `MusicalInstrument` linaloitwa `instrument_1`.

```js
({
    test: () => runPython(`
    assert isinstance(instrument_1, MusicalInstrument)
    `)
})
```

Unapaswa kupitisha `Oboe` na `woodwind` kama hoja wakati wa kuunda `instrument_1`.

```js
({ 
    test: () => runPython(`
        assert instrument_1.name == "Oboe"
        assert instrument_1.instrument_type == "woodwind" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--
```
