---
id: 6895dffc64e97df0e62e9ef9
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Njia ni kitendakazi kilichofafanuliwa ndani ya darasa la vitu ambacho kinaweza kutekeleza vitendo kwa kutumia data ya darasa hilo. Hapa ni jinsi unavyounda njia:

```python
def method_name(parameter1, parameter2,...):
    # Method contents will go here
```

Katika Python, njia ya `__init__` ni njia maalum inayojulikana kama njia ya kuanzisha, ambayo huitwa unapounda mfano mpya wa darasa.

Ndani ya darasa lako la `MusicalInstrument`, ongeza njia ya `__init__` yenye vigezo vitatu: `self`, `name`, na `instrument_type`.

Kigezo cha `self` ni njia ya kawaida ya kurejelea mfano wa darasa na kinahitajika kama kigezo cha kwanza kwa njia zote za mfano ndani ya darasa.

# --hints--

Unapaswa kufafanua njia ya `__init__` ndani ya darasa lako la `MusicalInstrument`.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").has_function("__init__")`))
});
```

Njia yako ya `__init__` inapaswa kuwa na vigezo vitatu: `self`, `name`, na `instrument_type`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").has_args("self, name, instrument_type")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
--fcc-editable-region--
    
    pass
--fcc-editable-region--
```
