---
id: 6895e197180c9af6988e1e19
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Tengeneza mfano wa pili uitwao `instrument_2` kwa ajili ya `Trumpet` ambayo ni chombo cha `brass`.

Angalia jinsi unavyoweza kutengeneza mifano mingi kutoka kwenye kiolezo kimoja cha darasa la vitu. Kila mfano huhifadhi thamani zake za sifa za kipengele kwa uhuru kutoka kwa mifano mingine.

# --hints--

Unapaswa kutengeneza mfano uitwao `instrument_2` wa darasa la vitu la `MusicalInstrument`.

```js
({
    test: () => runPython(`
    assert isinstance(instrument_2, MusicalInstrument)
    `)
})
```

Unapaswa kupitisha `Trumpet` na `brass` kama hoja wakati wa kutengeneza `instrument_2`.

```js
({ 
    test: () => runPython(`
        assert instrument_2.name == "Trumpet"
        assert instrument_2.instrument_type == "brass" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
--fcc-editable-region--

--fcc-editable-region--
```
