---
id: 6895e4e023e07603bcef46f9
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Sasa tumia njia ulizotengeneza.

Fanya wito wa njia ya `play()` kwenye mfano wako wa `instrument_1`. Hii itachapisha ujumbe kuhusu kucheza Oboe.

Ili kufanya wito wa njia kwenye mfano, tumia uandishi wa nukta:

```python
instance_name.method_name()
```

# --hints--

Unapaswa kufanya wito wa njia ya `play()` kwenye mfano wa `instrument_1`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_stmt("instrument_1.play()")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--

--fcc-editable-region--
```
