---
id: 6895e3f91c069e0068a42582
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

Ndani ya njia yako ya `play`, tumia `print()` kuonyesha ujumbe kuhusu kupiga chombo.

Tumia mfuatano wa herufi wa f-string kuingiza jina la chombo katika ujumbe: `The [instrument name] is fun to play!`. `[instrument name]` inapaswa kubadilishwa na jina halisi la chombo.

# --hints--

Njia yako ya `play` inapaswa kutumia mfuatano wa herufi wa f-string kuonyesha ujumbe `The [instrument name] is fun to play!`. Tumia `self.name` kupata jina la chombo.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_call('print(f"The {self.name} is fun to play!")')`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        pass
--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
