---
id: 6895e2becb7890fad818ad4d
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Sasa kwa kuwa umeelewa jinsi ya kupata sifa za kipengele cha kitu, ondoa wito zilizopo za `print`.

# --hints--

Unapaswa kuondoa tamko la `print(instrument_1.name)`.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

Unapaswa kuondoa tamko la `print(instrument_1.instrument_type)`.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

Unapaswa kuondoa tamko la `print(instrument_2.name)`.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

Unapaswa kuondoa tamko la `print(instrument_2.instrument_type)`.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
print(instrument_1.name)
print(instrument_1.instrument_type)

print(instrument_2.name)
print(instrument_2.instrument_type)
--fcc-editable-region--
```
