---
id: 6895e4940e39950288e406c2
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Ifuatayo, ongeza njia nyingine katika darasa la vitu lako liitwalo `get_fact` ambayo inarudisha ukweli kuhusu chombo.

Njia hii inapaswa kuwa na kigezo `self` pekee na irudishe mfuatano wa herufi wa aina ya f-string unaosema `The [instrument name] is part of the [instrument type] family of instruments.`

Tofauti na njia ya `play` ambayo huchapisha ujumbe moja kwa moja, njia hii inapaswa kurudisha mfuatano wa herufi ambao unaweza kutumika mahali pengine katika msimbo wako.

# --hints--

Unapaswa kufafanua njia liitwalo `get_fact` yenye kigezo `self`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").has_args("self")`))
});
```

Njia yako ya `get_fact` inapaswa kurudisha f-string inayojumuisha jina na aina ya chombo.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").find_return().is_equivalent("return f'The {self.name} is part of the {self.instrument_type} family of instruments.'")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

--fcc-editable-region--

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')
```
