---
id: 6895e235935106f88281b648
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Unaweza kupata sifa za mfano wako kwa kutumia uandishi wa nukta:

```py
instance_name.attribute_name
```

Hii inakuwezesha kupata thamani zilizohifadhiwa wakati wa uanzishaji.

Jaribu kuchapisha sifa za `name` na `instrument_type` za mifano yako yote miwili. Utaweza kuona thamani ulizoweka wakati wa kuziunda.

# --hints--

Unapaswa kuchapisha sifa ya `name` ya `instrument_1`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

Unapaswa kuchapisha sifa ya `instrument_type` ya `instrument_1`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

Unapaswa kuchapisha sifa ya `name` ya `instrument_2`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

Unapaswa kuchapisha sifa ya `instrument_type` ya `instrument_2`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
```
