---
id: 6895e0452b3370f26f365a46
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Katika njia ya `__init__`, unahitaji kuweka thamani za vigezo katika sifa za kipengele cha mfano. Sifa za kipengele cha mfano ni vigezo vinavyomilikiwa na mfano maalum wa darasa la vitu.

Ndani ya njia yako ya `__init__`, weka thamani ya kigezo cha `name` katika `self.name` na thamani ya kigezo cha `instrument_type` katika `self.instrument_type`.

# --hints--

Unapaswa kuweka thamani ya kigezo cha `name` katika `self.name` ndani ya njia ya `__init__`.

```js
 ({
    test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[0].is_equivalent("self.name = name")`))
});
```

Unapaswa kuweka thamani ya kigezo cha `instrument_type` katika `self.instrument_type` ndani ya njia ya `__init__`.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[1].is_equivalent("self.instrument_type = instrument_type")`))
})

```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
--fcc-editable-region--
        pass
--fcc-editable-region--
```
