---
id: 6895e5337b1914050b0af7bd
title: Hatua 12
challengeType: 20
dashedName: step-12
---

# --description--

Sasa, fanya wito wa njia ya `get_fact()` kwenye mfano wa `instrument_1` na chapisha matokeo.

Kwa kuwa `get_fact()` inarudisha mfuatano wa herufi badala ya kuuchapisha moja kwa moja, unahitaji kutumia `print()` kuonyesha mfuatano wa herufi uliorudishwa.

# --hints--

Unapaswa kufanya wito wa `get_fact()` kwenye `instrument_1` na chapisha matokeo.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.get_fact())")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


--fcc-editable-region--
instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

instrument_1.play()

--fcc-editable-region--
```
