---
id: 6895e356fd877efd6e3657ee
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--


Ndani ya darasa la `MusicalInstrument`, fafanua njia nyingine iitwayo `play` yenye kigezo kimoja tu cha `self`. Kigezo cha `self` kinaruhusu njia kupata sifa za kipengele za mfano maalum.

# --hints--

Unapaswa kufafanua njia iitwayo `play` yenye kigezo cha `self`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_args("self")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--
    
--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')
```
