---
id: a105e963526e7de52b219be9
title: "implementare un'unione ordinata unica"
challengeType: 26
dashedName: implement-a-unique-sorted-union
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione chiamata `uniteUnique`.
1. La funzione `uniteUnique` dovrebbe accettare due o più array come argomenti.
1. La funzione dovrebbe restituire un nuovo array che contiene valori unici presi dagli array argomento, nell'ordine in cui appaiono per la prima volta negli argomenti. Per esempio, un input come `[1, 2, 4], [2, 3, 5]` dovrebbe produrre un output come `[1, 2, 4, 3, 5]`.

# --hints--

Dovresti avere una funzione `uniteUnique`.

```js
assert.isFunction(uniteUnique);
```

`uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1])` dovrebbe restituire `[1, 3, 2, 5, 4]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

`uniteUnique([1, 2, 3], [5, 2, 1])` dovrebbe restituire `[1, 2, 3, 5]`.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1]), [1, 2, 3, 5]);
```

`uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8])` dovrebbe restituire `[1, 2, 3, 5, 4, 6, 7, 8]`.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8]), [
  1,
  2,
  3,
  5,
  4,
  6,
  7,
  8
]);
```

`uniteUnique([1, 3, 2], [5, 4], [5, 6])` dovrebbe restituire `[1, 3, 2, 5, 4, 6]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 4], [5, 6]), [1, 3, 2, 5, 4, 6]);
```

`uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1])` dovrebbe restituire `[1, 3, 2, 5, 4]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const uniteUnique = (...arrays) => {
  return arrays.reduce((a, b) => {
    return [
      ...a,
      ...b.filter((e, currentIndex) => b.indexOf(e) === currentIndex && a.indexOf(e) === -1)
    ];
  }, []);
};
```
