---
id: 69aaafeaa4037197c75587cd
title: تنفيذ طابور
challengeType: 26
dashedName: implement-a-queue
---

# --description--

في هذا المختبر، ستنفذ هيكل بيانات الطابور باستخدام الدوال. الطابور هو هيكل بيانات من نوع الأول يدخل أولاً يخرج أولاً (FIFO) حيث تُضاف العناصر إلى نهاية الطابور وتُزال من بدايته.

**الهدف:** حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك دالة `enqueue` تضيف عنصرًا إلى نهاية الطابور.
1. يجب أن يكون لديك دالة `dequeue` تزيل وتُرجع العنصر الموجود في بداية الطابور.
1. يجب أن يكون لديك دالة `front` تُرجع العنصر الموجود في بداية الطابور دون إزالته.
1. يجب أن يكون لديك دالة `size` تُرجع عدد العناصر في الطابور.
1. يجب أن يكون لديك دالة `isEmpty` تُرجع `true` إذا كان الطابور فارغًا و`false` خلاف ذلك.

**ملاحظة:** تعتمد معظم الاختبارات على دالة `enqueue`. نفذها أولاً، إذ تتطلب اختبارات `dequeue` و`front` و`size` و`isEmpty` إضافة عناصر إلى الطابور.

# --hints--

يجب أن يكون لديك دالة `enqueue`.

```js
assert.isFunction(enqueue);
```

يجب أن تضيف دالة `enqueue` عنصرًا إلى نهاية الطابور.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

يجب أن يكون لديك دالة `dequeue`.

```js
assert.isFunction(dequeue);
```

يجب أن تزيل دالة `dequeue` العنصر الموجود في بداية الطابور وتُرجعه.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(dequeue(queue), 'first');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(size(queue), 1);
```

يجب أن تُرجع دالة `dequeue` القيمة `undefined` عند استدعائها على طابور فارغ.

```js
const queue = initQueue();
assert.isUndefined(dequeue(queue));
```

يجب أن تستمر الدالة `enqueue` في إضافة العناصر إلى نهاية الطابور بعد عمليات الإزالة.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(dequeue(queue), 'first');
enqueue(queue, 'third');
assert.strictEqual(front(queue), 'second');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(dequeue(queue), 'third');
```

يجب ألا تشترك نسخ الطابور المختلفة في الحالة أو البيانات.

```js
const queueOne = initQueue();
const queueTwo = initQueue();
enqueue(queueOne, 'first');
enqueue(queueTwo, 'second');
assert.strictEqual(front(queueOne), 'first');
assert.strictEqual(front(queueTwo), 'second');
assert.strictEqual(size(queueOne), 1);
assert.strictEqual(size(queueTwo), 1);
```

يجب أن يكون لديك دالة `front`.

```js
assert.isFunction(front);
```

يجب أن تُرجع دالة `front` العنصر الموجود في بداية الطابور دون إزالته.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

يجب أن تُرجع دالة `front` القيمة `undefined` عند استدعائها على طابور فارغ.

```js
const queue = initQueue();
assert.isUndefined(front(queue));
```

يجب أن يكون لديك دالة `size`.

```js
assert.isFunction(size);
```

يجب أن تُرجع دالة `size` العدد الصحيح للعناصر في الطابور.

```js
const queue = initQueue();
assert.strictEqual(size(queue), 0);
enqueue(queue, 'first');
assert.strictEqual(size(queue), 1);
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(size(queue), 3);
dequeue(queue);
assert.strictEqual(size(queue), 2);
```

يجب أن يكون لديك دالة `isEmpty`.

```js
assert.isFunction(isEmpty);
```

يجب أن تُرجع دالة `isEmpty` القيمة `true` للطابور الفارغ.

```js
const queue = initQueue();
assert.isTrue(isEmpty(queue));
```

يجب أن تُرجع دالة `isEmpty` القيمة `false` للطابور غير الفارغ.

```js
const queue = initQueue();
enqueue(queue, 'element');
assert.isFalse(isEmpty(queue));
```

يجب أن تُرجع دالة `isEmpty` القيمة `true` بعد إزالة جميع العناصر من الطابور.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
dequeue(queue);
dequeue(queue);
assert.isTrue(isEmpty(queue));
```

# --seed--

## --seed-contents--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  
}

function dequeue(queue) {
  
}

function front(queue) {
  
}

function size(queue) {
  
}

function isEmpty(queue) {
  
}
```

# --solutions--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  queue.collection.push(element);
}

function dequeue(queue) {
  return queue.collection.shift();
}

function front(queue) {
  return queue.collection[0];
}

function size(queue) {
  return queue.collection.length;
}

function isEmpty(queue) {
  return queue.collection.length === 0;
}
```
