---
id: 669fcb06c3034a39f5431a38
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

To prepare creating some actual content, add a `body` element below the `head` element.

# --hints--

يجب أن يكون لديك العلامة المفتوحة (opening tag) الآتية `<body>`.

```js
assert.match(code, /<body>/i);
```

يجب أن يكون لديك العلامة المغلقة (closing tag) الآتي `</body>`.

```js
assert.match(code, /<\/body>/i);
```

لا يجب عليك تغيير عنصر `head`. تأكد من أنك لم تقم بحذف علامة الإغلاق الخاصة بك.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

عنصر `body` الخاص بك يجب أن يأتي بعد عنصر `head` الخاص بك.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>

  --fcc-editable-region--
</html>
```
