---
id: 68551b677a5f6898af364b85
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop lernen Sie mehr über Tailwind CSS, indem Sie dessen Flexbox-Utility-Klassen verwenden, um eine responsive Call-to-Action-(CTA-)Komponente zu erstellen.

Das grundlegende HTML-Boilerplate wurde für Sie bereitgestellt, sodass Sie zuerst die HTML-Elemente erstellen müssen.

Beginnen Sie damit, ein `div`-Element zu erstellen. Erstellen Sie innerhalb dieses `div` zwei weitere separate `div`-Elemente.

# --hints--

Sie sollten ein `div`-Element erstellen.

```js
assert.exists(document.querySelector("div"));
```

Sie sollten zwei weitere `div`-Elemente innerhalb des `div`-Elements haben.

```js
assert.lengthOf(document.querySelectorAll("div > div"), 2)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
