---
id: 6855255d96c1441bd919dc5d
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Wechseln Sie zum `span`-Element mit dem Text `Soundflow`. Der Text darin muss in Großbuchstaben sein.

Hier ist ein Beispiel, das zeigt, wie Sie die Utility-Klasse `uppercase` verwenden:

```html
<p class="uppercase">Quincy Larson</p>
```

Fügen Sie dem `span`-Element die Klasse `uppercase` hinzu.

# --hints--

Ihr `span`-Element sollte die Utility-Klasse `uppercase` haben.

```js
const spanEl = document.querySelector("span");
assert.isTrue(spanEl.classList.contains("uppercase"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        --fcc-editable-region--
        <span>Soundflow</span>
        --fcc-editable-region--
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
