---
id: 685523c05fb41d0848c4b52f
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Innerhalb des ersten `div`-Elements erstellen Sie ein `span`-Element mit dem Text `Soundflow`, ein `h1`-Element mit dem Text `Discover New Music` und ein `p`-Element mit dem Text `Stream your favorite tracks and discover new artists.`

# --hints--

Sie sollten ein `span`-Element mit dem Text `Soundflow` innerhalb des ersten `div`-Elements erstellen.

```js
assert.equal(document.querySelector("div > div:first-child > span")?.textContent, "Soundflow")
```

Sie sollten ein `h1`-Element mit dem Text `Discover New Music` innerhalb des ersten `div`-Elements erstellen.

```js
assert.equal(document.querySelector("div > div:first-child > h1")?.textContent, "Discover New Music")
```

Sie sollten ein `p`-Element mit dem Text `Stream your favorite tracks and discover new artists.` innerhalb des ersten `div`-Elements erstellen.

```js
assert.equal(document.querySelector("div > div:first-child > p")?.textContent, "Stream your favorite tracks and discover new artists.")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </div>
      <div></div>
    </div>
  </body>
</html>
```
