---
id: 68552464ef378e1051a55c28
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Erstellen Sie innerhalb des zweiten `div` zwei Ankerelemente. Das erste sollte den Text `Learn more` haben, und das zweite den Text `Start listening`.

Für das `href`-Attribut der beiden Ankerelemente setzen Sie beide auf `#`.

# --hints--

Sie sollten zwei `a`-Elemente innerhalb des zweiten `div`-Elements erstellen

```js
assert.lengthOf(document.querySelectorAll("div > div:nth-child(2) > a"), 2)
```

Ihr erstes `a`-Element sollte den Text `Learn more` und ein `href`-Attribut mit dem Wert `#` innerhalb des zweiten `div`-Elements haben.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[0]?.textContent, "Learn more")
assert.equal(links[0]?.getAttribute("href"), "#")
```

Ihr zweites `a`-Element sollte den Text `Start listening` und ein `href`-Attribut mit dem Wert `#` innerhalb des zweiten `div`-Elements haben.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[1]?.textContent, "Start listening")
assert.equal(links[1]?.getAttribute("href"), "#")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
