---
id: 6858e6408a0999603b2a4fee
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

Fügen Sie für das `h1` die Klassen `font-bold` und `text-4xl` hinzu. Dadurch wird der Text fetter, größer und erhält außerdem mehr Zeilenabstand.

Fügen Sie außerdem die Klasse `my-4` hinzu, um den oberen und unteren Außenabstand des `h1` um 16 Pixel, also `1rem`, zu vergrößern.

# --hints--

Ihr `h1`-Element sollte die Utility-Klasse `font-bold` haben.

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("font-bold"));
```

Ihr `h1`-Element sollte die Utility-Klasse `text-4xl` haben.

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("text-4xl"));
```

Ihr `h1`-Element sollte die Utility-Klasse `my-4` haben.

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("my-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        --fcc-editable-region--
        <h1>Discover New Music</h1>
        --fcc-editable-region--
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
