---
id: 6858f3a67aa8580ff37016c7
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Wenn Sie die Breite des Vorschau-Tabs wieder auf die Breite eines mobilen Geräts einstellen, werden Sie feststellen, dass der Inhalt immer nebeneinander angeordnet ist. Das ist auf mobilen Geräten nicht erwünscht.

Um das zu beheben, fügen Sie der `flex-col` die Klasse `div` hinzu. Dadurch ändert sich die Flex-Richtung zu Spalte, sodass die Elemente übereinander gestapelt werden.

So verwenden Sie die Utility-Klasse `flex-col`:

```html
<div class="flex flex-col">
  <element></element>
  <element></element>
</div>
```

# --hints--

Ihr öffnendes `div`-Element sollte die Utility-Klasse `flex-col` enthalten.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex-col"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
