---
id: 6858eadd3144c89d45e02bc7
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Wenn Sie genau hinsehen, ist das äußere `div` zu nah am oberen Rand und der Inhalt darin ist zu sehr nach links gedrängt. Deshalb sollten Sie jetzt an etwas Abstand arbeiten.

Fügen Sie die Klasse `mt-8` zu `div` hinzu, um es etwas nach unten zu verschieben, und `p-4`, um den Text vom linken Rand wegzurücken.

# --hints--

Ihr äußeres `div`-Element sollte die Utility-Klasse `mt-8` haben.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mt-8"));
```

Ihr äußeres `div`-Element sollte die Utility-Klasse `p-4` haben.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white">
--fcc-editable-region--
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
