---
id: 6858f90c54647a48ddb8822d
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

An diesem Punkt werden die Inhalte, also die beiden `div`-Elemente innerhalb des äußeren `div`, unabhängig von der Bildschirmgröße übereinander gestapelt.

Das liegt daran, dass Sie die `flex-col`-Klasse im umschließenden `div`-Element haben. Um das zu beheben, verwenden Sie das `lg:`-Präfix, um die Klasse `flex-row` dem `div` hinzuzufügen.

Das bedeutet, dass auf Geräten mit größeren Bildschirmen die Flex-Richtung `row` sein wird. Um das wieder zu sehen, passen Sie die Browserbreite auf einen größeren Wert an.

# --hints--

Ihr äußeres `div`-Element sollte die Utility-Klasse `lg:flex-row` haben.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("lg:flex-row"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
