---
id: 6858ee73ed8106ccb214942a
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Da jetzt alles links ausgerichtet ist, müssen Sie den gesamten Inhalt zentrieren, indem Sie einen Außenabstand von auto auf der x-Achse hinzufügen.

Fügen Sie die Tailwind-Klasse, die dem entspricht, dem äußeren `div`-Element hinzu. Um diesen Effekt zu sehen, passen Sie den Vorschaureiter so an, dass der gesamte Inhalt zentriert ist.

# --hints--

Ihr äußeres `div`-Element sollte die Utility-Klasse `mx-auto` haben.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2">
--fcc-editable-region--
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
