---
id: 6858ebf2552046ab3de9f17c
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

An diesem Punkt werden Sie feststellen, dass rechts vom Inhalt im `div` zu viel Platz ist. Dies können Sie durch Hinzufügen einer Breite minimieren.

Fügen Sie der `md:w-1/2` die Klasse `div` hinzu, um die Breite auf mittelgroßen und größeren Bildschirmen auf 50 % zu reduzieren. Das sind Bildschirme mit einer Breite von 768 Pixeln und mehr.

Um das in Aktion zu sehen, erweitern Sie den Vorschau-Tab.

# --hints--

Ihr äußeres `div`-Element sollte die Utility-Klasse `md:w-1/2` haben.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("md:w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
