---
id: 6858ef974c8ad5dcd7abb9f1
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Der Inhalt innerhalb von `div` kann besser verteilt und angeordnet werden. Für die restliche Verteilung und Anordnung müssen Sie Flexbox verwenden.

Fügen Sie der `flex` die Klasse `div` hinzu, damit die beiden untergeordneten `div`-Elemente nebeneinander sitzen.

Fügen Sie außerdem `justify-around` hinzu, damit der Abstand um und zwischen dem Inhalt gleichmäßig verteilt ist, und `items-center`, damit der gesamte Inhalt schließlich im `div` zentriert ist.

# --hints--

Ihr äußeres `div`-Element sollte die Utility-Klasse `flex` haben.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex"));
```

Ihr äußeres `div`-Element sollte die Utility-Klasse `justify-around` haben.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("justify-around"));
```

Ihr äußeres `div`-Element sollte die Utility-Klasse `items-center` haben.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto">
--fcc-editable-region--
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
