---
id: 6858e9f87e4faa903267d810
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Da der Text im äußeren `div` standardmäßig schwarz ist, ist er nicht gut genug sichtbar. Fügen Sie daher der `text-white` die Klasse `div` hinzu.

# --hints--

Ihr äußeres `div`-Element sollte die Utility-Klasse `text-white` haben.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("text-white"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600">
--fcc-editable-region--
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
