---
id: 685972d14c6d8182fbae71f1
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Um die Ankerelemente mehr wie Schaltflächen aussehen zu lassen, fügen Sie jedem von ihnen die Klassen `px-4` und `py-2` hinzu. `px-4` erhöht den Innenabstand rechts und links um `1rem`, und `py-2` erhöht den Innenabstand oben und unten um `0.5rem`.

# --hints--

Ihr erstes Ankerelement sollte die Utility-Klasse `px-4` haben.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("px-4"));
```

Ihr erstes Ankerelement sollte die Utility-Klasse `py-2` haben.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("py-2"));
```

Ihr zweites Ankerelement sollte die Utility-Klasse `px-4` haben.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("px-4"));
```

Ihr zweites Ankerelement sollte die Utility-Klasse `py-2` haben.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("py-2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        <a
          href="#"
          class="bg-white hover:bg-gray-100 text-indigo-600"
        >
          Learn more
        </a>
        <a
          href="#"
          class="bg-pink-500 hover:bg-pink-600 text-white"
         >
          Start listening
        </a>
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
