---
id: 68597114730d7c6b8838d204
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

Der Absatz unter dem `h1` sieht gut aus, daher benötigt er keine weitere Gestaltung. Das bedeutet, Sie können sich den Ankerelementen im zweiten `div` Element zuwenden.

Geben Sie dem ersten Ankerelement die Klassen `bg-white`, `hover:bg-gray-100` und `text-indigo-600`. Dadurch erhält es einen weißen Hintergrund, eine Textfarbe in Indigo und einen leicht helleren Hintergrund, wenn Sie mit der Maus darüber fahren.

Für das zweite Ankerelement geben Sie ihm die Klassen `bg-pink-500`, `hover:bg-pink-600` und `text-white`.

# --hints--

Ihr erstes Ankerelement sollte die Utility-Klasse `bg-white` haben.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("bg-white"));
```

Ihr erstes Ankerelement sollte die Utility-Klasse `hover:bg-gray-100` haben.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("hover:bg-gray-100"));
```

Ihr erstes Ankerelement sollte die Utility-Klasse `text-indigo-600` haben.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("text-indigo-600"));
```

Ihr zweites Ankerelement sollte die Utility-Klasse `bg-pink-500` haben.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("bg-pink-500"));
```

Ihr zweites Ankerelement sollte die Utility-Klasse `hover:bg-pink-600` haben.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("hover:bg-pink-600"));
```

Ihr zweites Ankerelement sollte die Utility-Klasse `text-white` haben.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("text-white"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        <a href="#">
          Learn more
        </a>
        <a href="#">
        --fcc-editable-region--
          Start listening
        </a>
      </div>
    </div>
  </body>
</html>
```
