---
id: 615f2d4150fe0d4cbd0f2628
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Debajo del elemento `h1`, agregue un elemento `p` con el texto `8 servings per container`.

# --hints--

Debe agregar un nuevo elemento `p`.

```js
assert.exists(document.querySelector('p'));
```

Su elemento `p` debe estar dentro de su elemento `body`.

```js
assert.equal(document.querySelector('p')?.parentElement?.localName, 'body');
```

Su elemento `p` debe venir después de su elemento `h1`.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.localName, 'h1');
```

Su elemento `p` debe tener el texto `8 servings per container`.

```js
assert.equal(document.querySelector('p')?.innerText, '8 servings per container');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
  --fcc-editable-region--
  
  --fcc-editable-region--
</body>
</html>
```

```css

```
