---
id: 67bf4350777ac6ffdc027745
title: カラーピッカーアプリを作成する
challengeType: 25
dashedName: build-a-color-picker-app
demoType: onClick
---

# --description--

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `ColorPicker` コンポーネントを定義してエクスポートする必要があります。
1. `useState` フックを使用する必要があります。
1. 白い背景の `#color-picker-container` 要素を持つ必要があります。
1. `#color-input` 要素は `color` 入力である必要があります。
1. `#color-input` は `#color-picker-container` の子要素である必要があります。
1. `#color-input` が変更されたとき、`#color-picker-container` の背景がその新しい値に設定される必要があります。

# --hints--

`ColorPicker` コンポーネントを定義してエクスポートしてください。

```js
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;

  const exports = {};
  const _a = eval(script);

  assert.property(exports, "ColorPicker");
```

`useState` フックを使用してください。

```js
    const abuseState = __helpers.spyOn(React, "useState");
    const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src === "index.jsx").innerText;
    const exports = {};
    const _a = eval(script);
    const _b = await __helpers.prepTestComponent(exports.ColorPicker);
    assert.isAtLeast(abuseState.calls.length, 1);
```

`useState` フックの戻り値を使用してください。

```js
assert.match(code, /=\s*(?:React\.)?useState\(/);
```

`useState` の戻り値を2つの変数に分割代入してください。

```js
assert.match(code, /\[.*,.*\]\s*=\s*(?:React\.)?useState\(/);
```

`#ffffff` を `useState` 呼び出しの初期値として渡してください。

```js
    const abuseState = __helpers.spyOn(React, "useState");
    const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src === "index.jsx").innerText;
    const exports = {};
    const _a = eval(script);
    const _b = await __helpers.prepTestComponent(exports.ColorPicker);
    assert.match(abuseState.calls[0], /#ffffff/i);
```

`color-picker-container` を `id` として持つ要素をレンダリングしてください。

```js
assert.exists(document.getElementById("color-picker-container"));
```

`#color-picker-container` は最初は白い背景である必要があります。

```js
const el = document.getElementById("color-picker-container");
const style = window.getComputedStyle(el);
assert.strictEqual(style.backgroundColor, "rgb(255, 255, 255)");
```

`input` 要素を `id` が `color-input` のものとしてレンダリングしてください。

```js
const inputMaybe = document.getElementById("color-input");
assert.exists(inputMaybe);
assert.strictEqual(inputMaybe?.tagName, "INPUT");
```

`#color-input` は `#color-picker-container` の子要素である必要があります。

```js
const inputPossibly = document.getElementById("color-input");
const parent = inputPossibly.parentElement;
assert.exists(parent);
assert.strictEqual(parent.getAttribute("id"), "color-picker-container");
```

`#color-input` は `color` 入力である必要があります。

```js
const inputDefinitely = document.getElementById("color-input");
assert.equal(inputDefinitely?.type, "color");
```

`#color-input` は `onChange` 属性を持つ必要があります。

```js
const definitelyAnInput = document.getElementById("color-input");
assert.exists(definitelyAnInput);
const key = Object.keys(definitelyAnInput).find(key => key.startsWith("__reactProps"));
const reactInstance = definitelyAnInput[key]
assert.exists(reactInstance);
assert.isFunction(reactInstance.onChange);
```

`#color-input` は `#ffffff` の初期値を持つ必要があります。

```js
const mostCertainlyAnInput = document.getElementById("color-input");
assert.exists(mostCertainlyAnInput);
assert.strictEqual(mostCertainlyAnInput.value, "#ffffff");
```

`#color-input` の初期値は `useState` フックで管理される必要があります。

```js
const abuseState = __helpers.spyOn(React, "useState");
const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src === "index.jsx").innerText;
const exports = {};
const _a = eval(script);

const rendered = await __helpers.prepTestComponent(exports.ColorPicker);
const inputElement = rendered.querySelector('#color-input');
assert.exists(inputElement, "Could not find #color-input element");

const initialValue = inputElement.value;

const stateIndex = abuseState.returns.findIndex(([stateValue]) => stateValue === initialValue);
assert.isAtLeast(stateIndex, 0, "The #color-input value must be managed by a useState hook");

const [, setterFunction] = abuseState.returns[stateIndex];

const newColor = '#123456';
await React.act(async () => {
  setterFunction(newColor);
});

assert.strictEqual(
  inputElement.value,
  newColor,
  "The #color-input value should update when the useState state changes"
);
```

`#color-input` の値が変更されたとき、`#color-picker-container` はその値を反映するように背景色を変更する必要があります。

```js
  const inputForSure = document.getElementById("color-input");
  assert.exists(inputForSure);
  const nativeValueSetter = Object.getOwnPropertyDescriptor(window.HTMLInputElement.prototype, "value").set;
  await React.act(async () => {
    nativeValueSetter.call(inputForSure, "#000001");
    inputForSure.dispatchEvent(new Event("input", { bubbles: true}));
  });
  const container = document.getElementById("color-picker-container");
  assert.exists(container);
  const style = window.getComputedStyle(container);
  assert.equal(style?.backgroundColor, "rgb(0, 0, 1)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Color Picker</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { ColorPicker } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<ColorPicker />);
    </script>
</body>

</html>
```

```css
body,
html {
    margin: 0;
    padding: 0;
    height: 100%;
    font-family: Arial, sans-serif;
}

#color-picker-container {
    height: 100vh;
    display: flex;
    justify-content: center;
    align-items: center;
    flex-direction: column;
    background-color: #ffffff;
}

input[type="color"] {
    position: absolute;
    margin-top: 50px;
    height: 40px;
}
```

```jsx
const { useState } = React;

export const ColorPicker = () => {};
```

# --solutions--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Color Picker</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { ColorPicker } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<ColorPicker />);
    </script>
</body>

</html>
```

```css
body,
html {
    margin: 0;
    padding: 0;
    height: 100%;
    font-family: Arial, sans-serif;
}

#color-picker-container {
    height: 100vh;
    display: flex;
    justify-content: center;
    align-items: center;
    flex-direction: column;
    background-color: #ffffff;
}

#pick-color-btn {
    padding: 15px 30px;
    font-size: 18px;
    color: white;
    background-color: #14183c;
    border: none;
    border-radius: 5px;
    cursor: pointer;
    margin-top: 20px;
}

#pick-color-btn:hover {
    background-color: #08131f;
}

input[type="color"] {
    position: absolute;
    margin-top: 50px;
    height: 40px;
}

.stylep{
    color: #14183c;
    font-size: 15px;
    font-weight: bold;
    margin-bottom: 50px;

}
```

```jsx
const { useState } = React;

export const ColorPicker = () => {
    const [bgColor, setBgColor] = useState('#ffffff');

    const handleColorChange = (event) => {
        setBgColor(event.target.value);
    };

    return (
            <div id="color-picker-container" style={{ backgroundColor: bgColor }}>
            <p className="stylep">
                Choose a color using the color input below:
            </p>

            <input
                id="color-input"
                type="color"
                value={bgColor}
                onChange={handleColorChange}
            />
        </div>
    );
}
```
