---
id: 67ef8f695db52583424bd118
title: Eine Seite für Filmkritiken gestalten
challengeType: 25
dashedName: design-a-movie-review-page
demoType: onClick
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten ein `main`-Element haben.
2. Im `main`-Element sollten Sie ein `h1`-Element für den Filmtitel haben.
3. Unter dem `h1`-Element sollten Sie ein `img`-Element einfügen, das das Filmcover anzeigt. Ihr `img`-Element sollte einen beschreibenden `alt`-Text enthalten, der das Bild beschreibt. Sie können gerne folgendes Bild verwenden: `https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png`.
4. Sie sollten ein `p`-Element mit einer kurzen Filmbeschreibung haben.
5. Sie sollten ein weiteres `p`-Element haben, um die Filmwertung anzuzeigen. Darin sollten Sie die folgenden Elemente in der angegebenen Reihenfolge einfügen:
   - Ein `b`-Element mit dem Text `Movie Rating:`.
   - Ein `span`-Element mit einem `aria-hidden`-Attribut, das auf `true` gesetzt ist, und einer visuellen Darstellung der Bewertung mit Sternen `⭐⭐⭐⭐⭐⭐⭐⭐⭐☆`.
   - Einen Zahlenwert, der die Bewertung darstellt, in runden Klammern (z. B. `9.2/10`) nach dem span.
6. Sie sollten ein `h2`-Element mit dem Text `Cast Members` haben.
7. Sie sollten ein `ul`-Element haben.
8. Im `ul`-Element sollten Sie mehrere `li`-Elemente haben, die jeweils ein `b`-Element für den Namen der Schauspielerin oder des Schauspielers enthalten, gefolgt vom entsprechenden Rollennamen, dem der Text `as` vorangestellt ist. (z. B. `James Holloway as Ethan Carter`).

# --hints--

Sie sollten ein `main`-Element haben.

```js
assert.isNotNull(document.querySelector("main"));
```

Innerhalb des `main`-Elements sollten Sie ein `h1`-Element mit dem Filmtitel haben.

```js
const h1 = document.querySelector("main > h1");
assert.isAbove(h1?.innerText.trim().length, 0);
```

Unterhalb des `h1`-Elements sollten Sie ein `img`-Element mit einem beschreibenden `alt`-Attribut haben.
  
```js
const img = document.querySelector("main > h1 + img");
assert.isNotNull(img);
assert.isNotEmpty(img.getAttribute("alt"));
```

Sie sollten ein `p`-Element nach dem Bild haben, das die Filmbeschreibung enthält.

```js
const description = document.querySelector("main img + p");
assert.isNotNull(description);
assert.isNotEmpty(description.innerText.trim());
```

Sie sollten ein weiteres `p`-Element haben, das die Bewertung enthält.

```js
const paragraphs = document.querySelectorAll("main p");
assert.isAtLeast(paragraphs.length, 2);
```

Innerhalb des zweiten `p`-Elements sollten Sie ein `b`-Element mit dem Text `Movie Rating: ` haben.

```js
const boldEl = document.querySelector("main p:nth-of-type(2) b");
assert.match(boldEl?.innerText.trim(), /^Movie Rating:?$/);
```

Es sollte ein `span`-Element innerhalb des Bewertungsabsatzes geben.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.isNotNull(spanEl);
```

Das `span`-Element innerhalb des Bewertungsabsatzes sollte zehn Sterne enthalten, entweder ausgefüllt (⭐) oder leer (☆), gefolgt von einem numerischen Wert in runden Klammern nach dem `span`.

```js  
const spanEl = document.querySelector("main p:nth-of-type(2) span");
const fullRatingText = document.querySelector("main p:nth-of-type(2)")?.textContent.replace(/\s+/g, ' ').trim();

assert.match(spanEl?.innerText.trim(), /^[⭐☆]{10}$/);
assert.match(fullRatingText, /Movie Rating:?\s*[⭐☆]{10}\s*\(\s*\d+(?:\.\d+)?\s*\/\s*10\s*\)$/);
```

Das `span`-Element innerhalb des Bewertungsabsatzes sollte ein `aria-hidden`-Attribut haben, das auf `true` gesetzt ist.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.strictEqual(spanEl?.getAttribute("aria-hidden"), "true");
```

Sie sollten ein `h2`-Element haben.

```js
assert.exists(document.querySelector("h2"));
```

Ihr `h2` sollte nach Ihrem zweiten `p`-Element kommen.

```js
const allParagraphs = document.querySelectorAll("p");
assert.isAtLeast(allParagraphs.length, 2);

const h2 = document.querySelector("h2");
const main = document.querySelector("main");

assert.exists(h2);
assert.exists(main);

const elements = Array.from(main.querySelectorAll("*"));
const lastPIndex = elements.indexOf(allParagraphs[1]);
const h2Index = elements.indexOf(h2);

assert.isAbove(h2Index, lastPIndex);
```

Ihr `h2` sollte den Text `Cast Members` haben.

```js
const h2 = document.querySelector("h2");
assert.equal(h2?.textContent.trim(), "Cast Members");
```

Sie sollten ein `ul`-Element nach dem `h2` haben.

```js
const ul = document.querySelector("main h2 + ul");
assert.isNotNull(ul);
```

Das `ul` sollte mehrere `li`-Elemente enthalten.

```js
const liItems = document.querySelectorAll("main ul li");
assert.isAbove(liItems.length, 1);
```

Jedes `li` sollte ein `b`-Element für den Namen des Schauspielers enthalten, gefolgt vom entsprechenden Charakternamen, dem der Text `as` vorangestellt ist. (z. B. `James Holloway as Ethan Carter`).

```js
const liItems = document.querySelectorAll("main ul li");  
assert.isNotEmpty(liItems);
liItems.forEach(li => {
  const bold = li.querySelector("b");
  assert.isNotNull(bold);
  const fullText = li.textContent.trim();
  const actorName = bold.textContent.trim();
  assert.isTrue(fullText.startsWith(actorName));
  const characterPart = fullText.slice(actorName.length).trim();
  assert.match(characterPart, /^as\s.+/);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
 <html lang="en">
 <head>
     <meta charset="UTF-8">
     <meta name="viewport" content="width=device-width, initial-scale=1.0">
     <title>Movie Review</title>
 </head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Movie Review</title>
  </head>
  <body>
    <main class="review">
      <h1>Movie Review: Rise Beyond</h1>
      <img
        src="https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png"
        alt="The Rise Beyond movie poster, starring Ethan Carter. A portrait of a young boy standing on a stage with his back toward the audience, holding a guitar at his side, staring at a backdrop of a cityscape with the sun rising behind its tall buildings."
      />
      <p>
        Rise Beyond is an uplifting drama that tells the story of Ethan Carter,
        a young man from a small town who dreams of becoming a world-class
        musician. Facing financial struggles and self-doubt, Ethan embarks on a
        journey of resilience, guided by his mentor Professor Adams and
        childhood friend Lena Mitchell. Through setbacks, failures, and moments
        of triumph, Rise Beyond reminds us that our greatest potential lies
        beyond our fears.
      </p>
      <p>
        <b>Movie Rating:</b>
        <span aria-hidden="true">⭐⭐⭐⭐⭐⭐⭐⭐⭐☆</span> (9.2/10)
      </p>

      <h2>Cast Members</h2>
      <ul>
        <li><b>James Holloway</b> as Ethan Carter</li>
        <li><b>Olivia Sterling</b> as Lena Mitchell</li>
        <li><b>William Lancaster</b> as Professor Adams</li>
        <li>
          <b>Maria Collins</b> as Evelyn Carter (Ethan’s Mother)
        </li>
      </ul>
    </main>
  </body>
</html>
```
