---
id: 60a3e3396c7b40068ad69977
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Використовуйте відступ, щоб налаштувати пробіли всередині елемента.

У `.frame` використайте скорочену властивість `padding`, щоб збільшити простір між елементами `.frame` і `.canvas` на `50px`. Ця скорочена властивість збільшить простір зверху, знизу, зліва та справа від кордону елемента і полотна всередині нього.

# --hints--

Ви повинні встановити властивість `padding` у значення `50px`.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

Ваш елемент `.frame` має мати значення `padding` рівне `50px`.

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
