---
id: 60a3e3396c7b40068ad69980
title: Крок 22
challengeType: 0
dashedName: step-22
---

# --description--

Додайте ще один `div` зі значенням `class` `two` безпосередньо під вашим елементом `one`. Це буде ваш другий прямокутник.

# --hints--

Ви не повинні змінювати існуючий елемент `.one`.

```js
assert(document.querySelectorAll('.one').length === 1);
```

У вашому елементі `div` має бути другий елемент `.canvas`.

```js
assert(document.querySelector('.canvas').children[1].tagName === 'DIV');
```

Другий елемент `div` повинен мати значення `class` `two`.

```js
assert(document.querySelector('.canvas').children[1].className.split(' ').includes('two'));
```

Ваш елемент `.two` має йти після елемента `.one`.

```js
assert(document.querySelector('.two').previousElementSibling.className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
