---
id: 60a3e3396c7b40068ad69974
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Кожна картина потребує рами.

Обгорніть елемент `.canvas` іншим `div`. Додайте цьому `div` клас `frame`.
# --hints--

Ви повинні додати новий елемент `div`.

```js
assert(document.querySelectorAll('div').length === 2)
```

Ваш елемент `.canvas` має бути вкладений у новий елемент `div`.

```js
assert(document.querySelector('.canvas').parentElement.tagName === 'DIV');
```

Ваш новий `div` має мати `class` зі значенням `frame`.

```js
assert(document.querySelector('div').className.split(' ').includes('frame'));
```

Ваш новий `div` має бути всередині вашого елемента `body`.

```js
assert(document.querySelector('div').parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--
  
  <div class="canvas">
  </div>
  
--fcc-editable-region--
  </body>
  ```
