---
id: 60a3e3396c7b40068ad6997e
title: Крок 20
challengeType: 0
dashedName: step-20
---

# --description--

Зараз `.one` центрований по горизонталі, але його верхнє поле виходить за межі полотна і заходить на кордон рамки, зсуваючи все полотно вниз на 20 пікселів.

Додайте `padding` зі значенням `1px` до елемента `.canvas`, щоб надати елементу `.one` щось твердое, від чого можна відштовхнутися.

# --hints--

Ви повинні встановити властивість `padding` у значення `1px`.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '1px');
assert(hasPadding);
```

Ваш елемент `.canvas` має мати значення `padding` рівне `1px`.

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('padding');
assert(canvasPadding === '1px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--
  
--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
