---
id: 60a3e3396c7b40068ad69979
title: Крок 15
challengeType: 0
dashedName: step-15
---

# --description--

Додайте новий елемент `div` всередину вашого елемента `.canvas`.

Надайте новому `div` атрибут `class` зі значенням `one`. Це буде ваш перший прямокутник.

# --hints--

Ви повинні створити новий елемент `div`.

```js
assert(document.querySelectorAll('div').length === 3);
```

Ви повинні вкласти новий елемент `div` всередину вашого елемента `.canvas`.

```js
assert(document.querySelector('.canvas').children[0].tagName === 'DIV');
```

Ваш новий `div` має мати атрибут `class` зі значенням `one`.

```js
assert(document.querySelector('.canvas').children[0].className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
