---
id: 60a3e3396c7b40068ad69978
title: Крок 14
challengeType: 0
dashedName: step-14
---

# --description--

Використовуйте поля, щоб налаштувати відступи поза межами елемента.

Використовуючи властивість `margin`, надайте елементу `.frame` вертикальне поле `20px` та горизонтальне поле `auto`. Це перемістить рамку вниз на 20 пікселів і горизонтально центровано на сторінці.

# --hints--

Ви повинні встановити властивість `margin` у значення `20px auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '20px auto');
assert(hasMargin);
```

Ваш елемент `.frame` має мати значення властивості `margin` рівне `20px auto`.

```js
const frameMargin = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('margin');
assert(frameMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
