---
id: 60a3e3396c7b40068ad69971
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Час для CSS.

Навіть якщо ваш `<div>` не має тексту, його все одно вважають блоком із вмістом.
Напишіть CSS-правило, яке використовує селектор класу `.canvas` і встановіть його властивість `width` у 500 пікселів.
Ось CSS-правило, яке встановлює ширину класу `card` у 300 пікселів:

```css
.card {
  width: 300px;
}
```

# --hints--

Ваш код має містити селектор `.canvas`.

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

Ви маєте встановити властивість `width` у значення `500px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

Ваш елемент `.canvas` має мати `width` `500px`.

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
