---
id: 60a3e3396c7b40068ad6996b
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Почніть свій проєкт, додавши елемент `img` з `src` рівним `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` всередині елемента `body`.

# --hints--

У вашому коді має бути елемент `img`.

```js
assert(document.querySelectorAll('img').length === 1);
```

Елемент `img` має мати `src` зі значенням `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
