---
id: 60a3e3396c7b40068ad69972
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Додайте властивість `height` зі значенням `600px` до вашого правила `.canvas`.

# --hints--

Ви повинні встановити властивість `height` у значення `600px`.

```js
const hasHeight = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.height === '600px');
assert(hasHeight);
```

Ваш елемент `.canvas` має мати ідентифікатор `height` зі значенням `600px`.

```js
const canvasHeight = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('height');
assert(canvasHeight === '600px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
