---
id: afcc8d540bea9ea2669306b6
title: Construye una función para repetir cadenas
challengeType: 26
dashedName: build-a-string-repeating-function
---

# --description--

En este laboratorio, crearás una función que repite una cadena dada un número específico de veces. Para el propósito de este laboratorio, *no* uses el método incorporado `.repeat()`.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes crear una función llamada `repeatStringNumTimes` que tome dos parámetros: una cadena y un número.
2. La función debe devolver la cadena repetida el número especificado de veces.
3. Si el número es menor o igual a cero, la función debe devolver una cadena vacía.

# --hints--

Debes crear una función llamada `repeatStringNumTimes`.

```js
assert.isFunction(repeatStringNumTimes);
```

`repeatStringNumTimes` debe tomar dos parámetros.

```js
assert.lengthOf(repeatStringNumTimes, 2);
```

La función `repeatStringNumTimes` siempre debe devolver una cadena.

```js
assert.isString(repeatStringNumTimes('hello', 3));
```

`repeatStringNumTimes("*", 3)` debe devolver la cadena `***`.

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

`repeatStringNumTimes("abc", 3)` debe devolver la cadena `abcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

`repeatStringNumTimes("abc", 4)` debe devolver la cadena `abcabcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

`repeatStringNumTimes("abc", 1)` debe devolver la cadena `abc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

`repeatStringNumTimes("*", 8)` debe devolver la cadena `********`.

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

`repeatStringNumTimes("abc", -2)` debe devolver una cadena vacía (`""`).

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

`repeatStringNumTimes("abc", 0)` debe devolver `""`.

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

El método incorporado `repeat()` no debe ser utilizado.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
