---
id: 64dcdee07a14f0299ec30559
title: 第 22 步
challengeType: 20
dashedName: step-22
---

# --description--

你撰寫了程式碼來尋找桿子之間允許的移動，但實際的移動可能是雙向的。  
在 `print()` 呼叫之後，宣告一個名為 `forward` 的變數並將其設為 `False`。你將使用該變數來檢查你需要將碟子在桿子之間移動的方向。

# --hints--

你應該宣告一個名為 `forward` 的變數。

```js
({ test: () => assert.match(code, /forward\s*=\s*./) })
```

你的 `forward` 變數應該設定為 `False`。

```js
({ test: () => assert.match(code, /forward\s*=\s*False/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
--fcc-editable-region--
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
--fcc-editable-region--
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
