---
id: 64dc976bf864693e668d67e8
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

現在呼叫你的函式（程式），並在終端機上查看輸出。

# --hints--

你應該呼叫你的 `move()` 函式（程式）。

```js
({test: () => assert.match(code, /^move\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    print(rods)
--fcc-editable-region--

--fcc-editable-region--
```
