---
id: 64dcd9bbc2268127e7898d77
title: 第 21 步
challengeType: 20
dashedName: step-21
---

# --description--

最後，當移動次數除以 3 沒有餘數時，允許在 `'B'` 和 `'C'` 之間移動。

為你的程式碼添加一個 `elif` 述語。然後，如果條件符合，列印適當的字串。 

# --hints--

當 `remainder == 0` 時，你應該有一個 `elif` 述語來執行。 

```js
({test: () => assert.match(code, /elif\s+remainder\s*==\s*0\s*:/)})
```

你應該列印字串 `f'Move {i + 1} allowed between {auxiliary} and {target}'`。 

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*auxiliary\s*\}\sand\s\{\s*target\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
