---
id: 64dc922df2919e3c38cead61
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

在你的程式碼頂端，宣告一個名為 `NUMBER_OF_DISKS` 的變數來儲存碟的數量，並將其值設為 `3`。然後，將傳入 `range()` 函式的第一個引數取代為你的新變數。

# --hints--

你應該宣告一個名為 `NUMBER_OF_DISKS` 的變數。

```js
({test: () => assert(__userGlobals.has('NUMBER_OF_DISKS')) })
```

`NUMBER_OF_DISKS` 變數應該有值 `3`。

```js
({test: () => assert.equal(__userGlobals.get('NUMBER_OF_DISKS'), 3) })
```

你應該將 `range()` 函式（程式）的第一個引數的值改成 `NUMBER_OF_DISKS` 變數。

```js
({test: () => assert.match(code, /('|")A\1\s*:\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

def move():
    pass
```
