---
id: 64dc92a9718fb53ca3d1bad7
title: 步驟 9
challengeType: 20
dashedName: step-9
---

# --description--

漢諾塔謎題只需 2<sup>n</sup> - 1 步即可解決，其中 n 是圓盤的數量。 聲明一個名爲 `number_of_moves` 的變量，並將總移動次數分配給該變量。

Python 中的冪運算符是 `**`。

# --hints--

你應該聲明一個名爲 `number_of_moves` 的變量。

```js
({ test: () => assert(__userGlobals.has('number_of_moves')) })
```

`number_of_moves` 的值應該是計算移動次數的表達式。

```js
({ test: () => {
    assert(runPython(`  
    number_of_moves == 2**NUMBER_OF_DISKS -1      
    `))
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
```
