---
id: 64dc8b1faf9c60304ca8a9b6
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個專案中，你將解決被稱為河內塔的數學謎題。該謎題由三根桿子和數個不同直徑的碟組成。

這個謎題的目標是將碟片從第一根桿子移動到第三根桿子，並遵循限制較大碟片放在較小碟片上方的特定規則。

首先建立一個名為 `rods` 的空字典來表述這些桿子。 

# --hints--

你應該有一個名為 `rods` 的變數。

```js
({ test: () => assert(__userGlobals.has('rods')) })
```

你的 `rods` 變數應該是一個空字典。

```js
({ test: () => assert(runPython(`
    rods == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
