---
id: 64dc922df2919e3c38cead61
title: 步驟 8
challengeType: 20
dashedName: step-8
---

# --description--

在代碼頂部，聲明一個名爲 `NUMBER_OF_DISKS` 的變量來存儲磁盤數量，並賦予其 `3` 的值。 然後，用新變量替換傳遞給 `range()` 函數的第一個參數。

# --hints--

你應該聲明一個名爲 `NUMBER_OF_DISKS` 的變量。

```js
({test: () => assert(__userGlobals.has('NUMBER_OF_DISKS')) })
```

`NUMBER_OF_DISKS` 變量的值應爲 `3`。

```js
({test: () => assert.equal(__userGlobals.get('NUMBER_OF_DISKS'), 3) })
```

你應該將 `range()` 函數的第一個參數的值更改爲 `NUMBER_OF_DISKS` 變量。

```js
({test: () => assert.match(code, /('|")A\1\s*:\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

def move():
    pass
```
