---
id: 64dc8ea01436383a88256d7f
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

謎題從第一根桿子上堆疊的碟開始，大小依序遞減，最小的碟在最上方，最大的碟在最下方。你需要用表述各種碟大小的數字來填充你的第一個列表。

與其手動將項目添加到第一個列表，不如使用 `range()` 函式（程式）產生一個從 `3` 到 `1` 的遞減數字序列，並將其指定給 `rods['A']`。這裡，`3` 表現堆疊底部最大的碟，`1` 表現堆疊頂部最小的碟。

語法是 `range(x, y, h)`，其中 `x` 是起始整數（包含），`y` 是最後一個整數（不包含），而 `h` 是序列中相鄰兩數字之間的差距。

# --hints--

你應該使用 `range()` 函式為 `rods['A']` 指定一個數字序列。呼叫 `range()` 函式的語法是 `range(x, y, h)`。

```js
({ test: () => assert(runPython(`
  type(rods['A']) is range
  `))
})
```

你應該使用 `range()` 函式（程式）將從 `3` 到 `1` 的數字序列指定給 `rods['A']`。

```js
({ test: () => assert(runPython(`
 rods['A'] == range(3, 0, -1)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': [],
    'B': [],
    'C': []
}
--fcc-editable-region--
```
