---
id: 64dc976bf864693e668d67e8
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

現在調用你的函數並查看終端上的輸出。

# --hints--

你應該調用你的 `move()` 函數。

```js
({test: () => assert.match(code, /^move\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    print(rods)
--fcc-editable-region--

--fcc-editable-region--
```
