---
id: 64dc963d866fff3dd0329993
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

列印你在前一步宣告的 `variable`，並隨意更改碟的數量，以查看所需的最少移動次數如何快速增加。

# --hints--

你應該列印 `number_of_moves` 變數。

```js
({ test: () => assert.match(code, /print\s*\(\s*number_of_moves\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move():
    pass
```
