---
id: 64dc916138dbe03bdfe6cec6
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

河內塔的目標是將所有的碟移動到最後一根桿子。為此，你必須遵守三個簡單的規則：

1. 你只能移動最上層的碟。
2. 你一次只能移動一個碟。
3. 你不能將較大的磁碟放在較小的磁碟上面

在你現有的程式碼下方，宣告一個名為 `move` 的空函式。稍後，你會使用該函式在桿子之間移動碟。現在，為避免錯誤，請在函式主體內使用 `pass` 關鍵字。

# --hints--

你應該宣告一個名為 `move` 的空函式（程式）。請記得在函式（程式）主體內使用帶有正確縮排的 `pass` 關鍵字。

```js
({ test: () => {
		assert(runPython(`
        import inspect
        inspect.isfunction(move)
    `))
	} 
})
```

# --seed--

## --seed-contents--

```py
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--
```
