---
id: 64dc9004f658183afa85f4f1
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

現在透過將 `'A'` 鍵傳入 `type()` 函式來檢查它的資料型別，並在終端機上列印結果。

# --hints--

你應該將你的 `'A'` 鍵傳遞給 `type()` 函式（程式）。

```js
({ test: () => assert.match(code, /type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)/) })
```

你應該使用 `print()` 函式列印 `rods['A']` 的型別。

```js
({ test: () => assert.match(code, /print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
--fcc-editable-region--
```
