---
id: 64dcd3d61c448e2676501f43
title: 步驟 20
challengeType: 20
dashedName: step-20
---

# --description--

當移動數值除以3的餘數等於2時，允許在`'A'`和`'B'`之間移動（源和備用棒）。

爲此添加 `elif` 語句。 然後，如果滿足條件，則打印適當的字符串。

# --hints--

你應該有一個 `elif` 語句在 `remainder == 2` 時執行。

```js
({test: () => assert.match(code, /elif\s+remainder\s*==\s*2:/)})
```

你應該打印字符串 `f'Move {i + 1} allowed between {source} and {auxiliary}'`。

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*source\s*\}\sand\s\{\s*auxiliary\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
