---
id: 64dc92a9718fb53ca3d1bad7
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

河內塔謎題可以在 2<sup>n</sup> - 1 步內解決，其中 n 是碟的數量。宣告一個名為 `number_of_moves` 的變數，並將總步數指定給這個變數。

Python 中的次方運算子是 `**`。

# --hints--

你應該宣告一個名為 `number_of_moves` 的變數。

```js
({ test: () => assert(__userGlobals.has('number_of_moves')) })
```

`number_of_moves` 的值應該是用來計算移動次數的表達式。

```js
({ test: () => {
    assert(runPython(`  
    number_of_moves == 2**NUMBER_OF_DISKS -1      
    `))
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
```
