---
id: 64dc985e6720a23edac63b51
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

允許在桿子之間的碟移動呈現每三次移動發生一次的重複樣式。 例如，在第 1、4 和 7 次移動時允許在桿子 `A` 和桿子 `C` 之間移動，該移動次數除以 3 的餘數相等於 1。

在先前創建的 `for` 迴圈內，將現有的 `print()` 呼叫取代為當 `(i + 1) % 3 == 1` 時觸發的 `if` 述語。在此 `if` 述語內，使用 f-string 列印 `f'Move {i + 1} allowed between {source} and {target}'`。請注意，`i + 1` 是移動次數，因為在第一次迭代時 `i` 為零。

# --hints--

你應該有一個 `if` 述語，當 `(i + 1) % 3 == 1` 時應該執行。 

```js
({test: () => assert.match(code, /if\s+\(\s*i\s*\+\s*1\s*\)\s*%\s*3\s*==\s*1\s*:/)})
```

你應該使用 f-string 列印字串 `f'Move {i + 1} allowed between {source} and {target}'`。

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*source\s*\}\sand\s\{\s*target\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        print(i)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
