---
id: 64dc90949ff85e3b37be40d0
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

`range()` 函式傳回一個不可變的數字序列。如你所見，`rods['A']` 的資料型別是 `range`，但你想要它是列表。

將你的 `range()` 呼叫傳遞給 `list()` 函式來執行該操作。

# --hints--

你應該將你的 `range()` 呼叫傳遞給 `list()` 函式（程式）。

```js
({ test: () => assert(runPython(`
  rods['A'] == list(range(3, 0, -1))
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
