---
id: 64dce7ee2494e82bf4205a3f
title: 第 24 步
challengeType: 20
dashedName: step-24
---

# --description--

另一種你必須將碟必然從 `source` 移動到 `target` 的情況是當 `source` 列表**不**為空**且** `source` 中的最後一個碟比 `target` 中的最後一個碟還要小。

為你的程式碼添加一個 `elif` 述語來檢查此條件。然後，如果條件符合，將 `forward` 變數設為 `True`。

# --hints--

你應該有一個 `elif` 述語來檢查 `rods[source]` 是否不為空且 `rods[source][-1] < rods[target][-1]`。

```js
const allowedRes = [
    "rods\\s*\\[\\s*source\\s*\\]",
    "len\\(rods\\s*\\[\\s*source\\s*\\]\\s*\\)\\s*>\\s*0",
    "len\\(rods\\s*\\[\\s*source\\s*\\]\\s*\\)\\s*>=\\s*1",
];
const re = new RegExp(`elif\\s+((${allowedRes.join(")|(")}))\\s+and\\s+(rods\\s*\\[\\s*source\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]\\s*<\\s*rods\\s*\\[\\s*target\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]|rods\\s*\\[\\s*target\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]\\s*>\\s*rods\\s*\\[\\s*source\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\])\\s*:`);
assert.match(code, re);
```

你應該在新的 `elif` 述語中將 `forward` 設定為 `True`。你不應該更改先前對 `forward` 的指派。

```js
const allowedRes = [
    "rods\\s*\\[\\s*source\\s*\\]",
    "len\\(rods\\s*\\[\\s*source\\s*\\]\\s*\\)\\s*>\\s*0",
    "len\\(rods\\s*\\[\\s*source\\s*\\]\\s*\\)\\s*>=\\s*1",
];
const re = new RegExp(`elif\\s+((${allowedRes.join(")|(")}))\\s+and\\s+(rods\\s*\\[\\s*source\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]\\s*<\\s*rods\\s*\\[\\s*target\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]|rods\\s*\\[\\s*target\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]\\s*>\\s*rods\\s*\\[\\s*source\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\])\\s*:\\s+forward\\s*=\\s*True`);
assert.match(code, re);
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            forward = False
--fcc-editable-region--            
            if not rods[target]:
                forward = True
--fcc-editable-region--
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
