---
id: 64dc97005bc7943e2978df0a
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

現在你可以移除你的 `print()` 呼叫。接著，在 `move()` 函式中，移除 `pass` 關鍵字並列印你的 `rods` 字典內容。

# --hints--

你的程式碼中不應該有 `print(number_of_moves)`。

```js
({ test: () => assert.isFalse(/print\s*\(\s*number_of_moves\s*\)/.test(code)) })
```

你應該移除 `pass` 關鍵字並列印 `rods` 字典。

```js
({ test: () => assert.match(code, /def\s+move\s*\(\s*\)\s*:\s+print\s*\(\s*rods\s*\)(?!\s*pass)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
print(number_of_moves)
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
--fcc-editable-region--
```
