---
id: 64dc90e837670a3ba26a1e61
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

既然型別已經是所需的 `list`，你可以移除 `print()` 呼叫。

# --hints--

你應該從你的程式碼中移除 `print(type(rods['A']))`。

```js
({ test: () => assert.match(code, /\}(?!\s+print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")A\1\s*\]\s*\)\s*\))/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
