---
id: 64ddd280d69fa818ca782d50
title: 第 29 步
challengeType: 20
dashedName: step-29
---

# --description--

如你所見，磁碟 1 每三次移動就來回移動一次。這是因為你仍需處理其他桿子之間的移動。  
與其重複你在前幾個步驟中撰寫並更改桿子的相同程式碼，不如將該程式碼移入一個函式中，以便在每個條件述語中呼叫。宣告一個名為 `make_allowed_move()` 的空函式，並且別忘了 `pass` 關鍵字。

# --hints--

你應該宣告一個名為 `make_allowed_move` 的空函式（程式）。請記得在函式（程式）主體內使用 `pass` 關鍵字並保持正確的縮排。

```js
({ test: () => assert(runPython(`
      import inspect
      inspect.isfunction(make_allowed_move)
  `)) 
})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            forward = False
            if not rods[target]:
                forward = True
            elif rods[source] and rods[source][-1] < rods[target][-1]:
                forward = True              
            if forward:
                print(f'Moving disk {rods[source][-1]} from {source} to {target}')
                rods[target].append(rods[source].pop())
            else:
                print(f'Moving disk {rods[target][-1]} from {target} to {source}')
                rods[source].append(rods[target].pop())
            
            # display our progress
            print(rods)
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
