---
id: 64de1469fccce61940f498c4
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

`rods` 字典將表述三根桿子及其上的硬碟。給它字串 `'A'`、`'B'` 和 `'C'` 作為鍵，並將它們各自設為空列表。

# --hints--

你的 `rods` 字典應該有一個 `'A'` 鍵。

```js
({ test: () => assert(runPython(`'A' in rods`)) })
```

`rods['A']` 應該是一個空的列表。

```js
({ test: () => assert(runPython(`
    rods['A'] == []
  `))
})
```

你的 `rods` 字典應該有一個 `'B'` 鍵。

```js
({ test: () => assert(runPython(`'B' in rods`)) })
```

`rods['B']` 應該是一個空的列表。

```js
({ test: () => assert(runPython(`
    rods['B'] == []
  `))
})
```

你的 `rods` 字典應該有一個 `'C'` 鍵。

```js
({ test: () => assert(runPython(`'C' in rods`)) })
```

`rods['C']` 應該是一個空的列表。

```js
({ test: () => assert(runPython(`
    rods['C'] == []
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {}
--fcc-editable-region--
```
