---
id: 64df3e2fac34d813d048f3f9
title: 步驟 43
challengeType: 20
dashedName: step-43
---

# --description--

你也不需要 `make_allowed_move` 和 `number_of_moves`。刪除整個函式和變數。

# --hints--

你應該刪除 `number_of_moves` 變數及其內容。

```js
({ test: () => assert.match(code, /NUMBER_OF_DISKS\s*=\s*4\s+(?=rods\s*=\s*\{)/) })
```

你應該刪除整個 `make_allowed_move` 函式（程式）。

```js
({ test: () => assert.match(code, /\}\s+(?=def\s+move\(\s*n\s*,\s*source\s*,\s*auxiliary\s*,\s*target\s*\)\s*:)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 4
number_of_moves = 2 ** NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def make_allowed_move(rod1, rod2):    
    forward = False
    if not rods[rod2]:
        forward = True
    elif rods[rod1] and rods[rod1][-1] < rods[rod2][-1]:
        forward = True              
    if forward:
        print(f'Moving disk {rods[rod1][-1]} from {rod1} to {rod2}')
        rods[rod2].append(rods[rod1].pop())
    else:
        print(f'Moving disk {rods[rod2][-1]} from {rod2} to {rod1}')
        rods[rod1].append(rods[rod2].pop())
    
    # display our progress
    print(rods, '\n')

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
