---
id: 64df496c6a8ddf1b38db1ed6
title: 步驟 50
challengeType: 20
dashedName: step-50
---

# --description--

在先前的步驟中，你撰寫了將子問題中最大的碟移動到標的桿的程式碼。

現在，你只需要為你已經移動的 `n - 1` 個碟添加另一個遞迴呼叫。複製第一個遞迴呼叫並將它貼到 `if` 區塊的結尾。

請注意，函式的引數順序不正確。試著找出正確的順序。

# --hints--

你應該在你的 `move` 函式結尾呼叫 `move(n - 1,  auxiliary, source, target)`。

```js
({ test: () => assert.match(code, /^\s{8}move\(\s*n\s*-\s*1\s*,\s*auxiliary\s*,\s*source\s*,\s*target\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
