---
id: 64df72c6e0183a191fcd72dc
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

在你的函式（程式）呼叫之前，寫一個註解說明 `initiate call from source A to target C with auxiliary B`。

# --hints--

你應該在你的函式呼叫之前撰寫給定的註解。

```js
({test: () => assert.match(code, /#\s*initiate\scall\sfrom\ssource\sA\sto\starget\sC\swith\sauxiliary\sB/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    print(rods)
--fcc-editable-region--

move(NUMBER_OF_DISKS, 'A', 'B', 'C')
--fcc-editable-region--
```
