---
id: 64df47b32b92301a815d5ef8
title: 步驟 49
challengeType: 20
dashedName: step-49
---

# --description--

起初，你剛剛添加的遞迴呼叫處理將 `n - 1` 個碟移動到第二根桿的子問題。

因此，`target` 引數對應你的第二根桿子，而 `auxiliary` 引數是第三根桿子。請記住，隨著遞迴進行，這些桿子會不斷交換。

修正遞迴呼叫中交換 `target` 和 `auxiliary` 的引數順序。

# --hints--

你應該修改 `move(n - 1, source, auxiliary, target)` 呼叫中引數的順序。

```js
({ test:  () => assert.match(code, /move\(\s*n\s*-\s*1\s*,\s*source\s*,\s*target\s*,\s*auxiliary\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, auxiliary, target)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
