---
id: 64df45a3ad4f8719e5355244
title: 步驟46
challengeType: 20
dashedName: step-46
---

# --description--

移動 `n - 1` 個圓盤的步驟可以進一步細分，直到只考慮單個圓盤。 這將是發生的第一步。 第一次移動發生後，通過展開遞歸調用生成後續移動。 請記住，在每個遞歸步驟中，杆所扮演的角色會在源、目標和輔助之間發生變化。

目前，每次遞歸調用都會打印 `rods` 字典，而不會對列表進行任何更改。 在 `print()` 調用之前，從 `rods[source]` 列表中刪除最後一個元素並將其附加到 `rods[target]` 列表中。

# --hints--

你應該從 `rods[source]` 列表中刪除最後一個元素，並在調用 `print` 之前將其附加到 `rods[target]` 列表中。

```js
({ test: () => assert.isTrue(runPython(`
_log = []
def capture_print(*args):
    _log.append(repr(args))

old_print = print
print = capture_print

old_rods = rods
rods2 = {
    'A': list(range(5, 0, -1)),
    'B': [],
    'C': []
}
rods = rods2

move(5, 'A', 'B', 'C')
rods = old_rods
print = old_print

_expected_prints = [
r"({'A': [5, 4, 3, 2], 'B': [], 'C': [1]}, '\\n')",
r"({'A': [5, 4, 3], 'B': [], 'C': [1, 2]}, '\\n')",
r"({'A': [5, 4], 'B': [], 'C': [1, 2, 3]}, '\\n')",
r"({'A': [5], 'B': [], 'C': [1, 2, 3, 4]}, '\\n')",
r"({'A': [], 'B': [], 'C': [1, 2, 3, 4, 5]}, '\\n')",
]

for args in _log:
    if not _expected_prints:
        break
    if args == _expected_prints[0]:
        _expected_prints.pop(0)

not rods2['A'] and rods2['C'] == list(range(1, 6)) and not _expected_prints
`))})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, auxiliary, target)

        # display starting configuration
        print(rods, '\n')

--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
