---
id: 64e33a488b4a2717fa22ebf3
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

在你的 `print()` 呼叫之前添加另一個註解，內容為 `display starting configuration`。

# --hints--

你應該使用提供的文字在你的 `print()` 呼叫之前添加註解。

```js
({ test: () => assert.match(code, /#\s*display\sstarting\sconfiguration\s+print\s*\(\s*rods\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    
    print(rods)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
