---
id: 64e34146860065146733883b
title: 步驟 53
challengeType: 20
dashedName: step-53
---

# --description--

雖然遞迴有時候較難理解，但它讓你能夠創建更簡潔的程式碼。在這種情況下，你甚至不需要區分碟的偶數和奇數。 

將 `NUMBER_OF_DISKS` 設定為 `5` 並檢查輸出。

# --hints--

你應該將 `NUMBER_OF_DISKS` 變數設定為 5。

```js
({test: () => assert.equal(__userGlobals.get('NUMBER_OF_DISKS'), 5) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 4
A = list(range(NUMBER_OF_DISKS, 0, -1))
B = []
C = []

def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        target.append(source.pop())
        
        # display our progress
        print(A, B, C, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, A, B, C)
--fcc-editable-region--
```
