---
id: 650c6082e5586f9e3acfcd3b
title: 步驟 52
challengeType: 20
dashedName: step-52
---

# --description--

現在，刪除 `rods` 字典，並將其鍵分別轉換成變數 `A`、`B` 和 `C`，同時保留它們的值。重構你的程式碼以反映這些變更。如果你在終端機上看到輸出，表示你正確完成了這個步驟。 

# --hints--

你應該將 `rods['A']` 轉換成名為 `A` 的變數並保持其值。

```js
assert.match(code, /A\s*=\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/);
```

你應該將 `rods['B']` 轉成名為 `B` 的變數並保持其值。

```js
assert.match(code, /B\s*=\s*\[\s*\]/);
```

你應該將 `rods['C']` 轉成名為 `C` 的變數並保持其值。

```js
assert.match(code, /C\s*=\s*\[\s*\]/);
```

你應該修改你的 `print()` 呼叫來列印 `A`、`B`、`C`，而不是 `rods` 物件。請在 `print()` 呼叫中保留換行字元。

```js
assert.match(code, /print\s*\(\s*A\s*,\s*B\s*,\s*C\s*,\s*('|")\\n\1\s*\)/);
```

你應該在你的程式碼中有 `target.append(source.pop())`。

```js
assert.match(code, /target\.append\s*\(\s*source\.pop\s*\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
--fcc-editable-region--
```

