---
id: 65782342850feb3b8d62f936
title: 步驟 18
challengeType: 20
dashedName: step-18
---

# --description--

由於你將多次使用表達式 `(i + 1) % 3`，因此將它儲存在變數中會比較方便。

在你的 `if` 述語正上方，宣告一個 `remainder` 變數並將值 `(i + 1) % 3` 指定給此變數。

# --hints--

你應該在 `if` 區塊正上方宣告一個 `remainder` 變數。

```js
({ test: () => assert.match(code, /^\s{8}remainder\s*=/m) })
```

你的 `remainder` 變數的值應該是表達式 `(i + 1) % 3`。

```js
({ test: () => assert.match(code, /^\s{8}remainder\s*=\s*\(\s*i\s*\+\s*1\s*\)\s*%\s*3/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        
        if (i + 1) % 3 == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
