---
id: 66fbcf750a62784cf11f5634
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

È il momento di testare ciascuna delle funzioni. Puoi usare una qualsiasi delle ricette per questo, ma questo tutorial inizierà con `recipe1`.

Crea due nuove variabili: `recipe1TotalIngredients` e `recipe1DifficultyLevel`. Assegna loro i valori chiamando la funzione corrispondente per ogni variabile e passando la proprietà `recipe1` appropriata.

Infine, stampa ciascuna variabile nella console per vedere i risultati.

# --hints--

Dovresti creare una variabile `recipe1TotalIngredients`.

```js
assert.isNotNull(recipe1TotalIngredients);
```

La tua variabile `recipe1TotalIngredients` dovrebbe essere impostata su `getTotalIngredients(recipe1.ingredients)`.

```js
assert.equal(recipe1TotalIngredients, getTotalIngredients(recipe1.ingredients));
```

Dovresti stampare `recipe1TotalIngredients` nella console.

```js
assert.match(code, /console\.log\(recipe1TotalIngredients\);?/);
```

Dovresti creare una variabile `recipe1DifficultyLevel`.

```js
assert.isNotNull(recipe1DifficultyLevel);
```

La tua variabile `recipe1DifficultyLevel` dovrebbe essere impostata su `getDifficultyLevel(recipe1.cookingTime)`.

```js
assert.equal(recipe1DifficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

Dovresti stampare `recipe1DifficultyLevel` nella console.

```js
assert.match(code, /console\.log\(recipe1DifficultyLevel\);?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
