---
id: 66fbcf750a62784cf11f5630
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Ora dovresti inserire i tre oggetti nell'array `recipes`. Per farlo, puoi usare il metodo `push()`.

Usa il metodo `push()` per inserire tutti gli oggetti ricetta nell'array `recipes`. Assicurati di inserire `recipe1`, `recipe2` e `recipe3` in questo ordine.

Elimina anche le variabili `recipe1Name`, `recipe2CookingTime` e `recipe3Ingredients` e le istruzioni `console.log` che stampano quelle variabili.

# --hints--

Dovresti inserire tutti gli oggetti ricetta nell'array `recipes`. Assicurati di inserirli nell'ordine in cui sono dichiarati.

```js
assert.deepEqual(recipes, [recipe1, recipe2, recipe3]);
```

Dovresti eliminare la variabile `recipe1Name`.

```js
assert.throws(() => recipe1Name, ReferenceError);
```

Dovresti eliminare l'istruzione `console.log(recipe1Name);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1Name\s*\)\s*;?/);
```

Dovresti eliminare la variabile `recipe2CookingTime`.

```js
assert.throws(() => recipe2CookingTime, ReferenceError);
```

Dovresti eliminare l'istruzione `console.log(recipe2CookingTime);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2CookingTime\s*\)\s*;?/);
```

Dovresti eliminare la variabile `recipe3Ingredients`.

```js
assert.throws(() => recipe3Ingredients, ReferenceError);
```

Dovresti eliminare l'istruzione `console.log(recipe3Ingredients);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe3Ingredients\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--
const recipe1Name = recipe1.name;
console.log(recipe1Name);

const recipe2CookingTime = recipe2.cookingTime;
console.log(recipe2CookingTime);

const recipe3Ingredients = recipe3.ingredients;
console.log(recipe3Ingredients);
--fcc-editable-region--
```
