---
id: 66fbcf750a62784cf11f562d
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Crea un oggetto chiamato `recipe1`. All'interno dell'oggetto `recipe1`, crea una proprietà `name` con valore `"Spaghetti Carbonara"`.

Sempre all'interno dell'oggetto `recipe1`, crea una proprietà `ingredients` con un array come valore. L'array deve contenere `spaghetti`, `Parmesan cheese`, `pancetta` e `black pepper`.

# --hints--

Dovresti creare un oggetto chiamato `recipe1`.

```js
assert.isObject(recipe1);
```

Il tuo oggetto `recipe1` dovrebbe avere una proprietà `name`.

```js
assert.property(recipe1, 'name');
```

La tua proprietà `name` dovrebbe essere impostata su `Spaghetti Carbonara`.

```js
assert.equal(recipe1.name, 'Spaghetti Carbonara');
```

Il tuo oggetto `recipe1` dovrebbe avere una proprietà `ingredients`.

```js
assert.property(recipe1, 'ingredients');
```

La tua proprietà `ingredients` dovrebbe avere come valore un array.

```js
assert.isArray(recipe1.ingredients);
```

L'array valore della tua proprietà `ingredients` dovrebbe contenere `spaghetti`, `Parmesan cheese`, `pancetta` e `black pepper`.

```js
assert.deepEqual(recipe1.ingredients, ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];
--fcc-editable-region--

--fcc-editable-region--
```
