---
id: 66fbcf750a62784cf11f5633
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

Crea una `getDifficultyLevel` funzione che prenda un numero che indica il tempo di cottura come parametro.

Se il tempo di cottura è minore o uguale a `30`, la funzione dovrebbe restituire `"easy"`. Se è minore o uguale a `60`, la funzione dovrebbe restituire `"medium"`. Altrimenti, la funzione dovrebbe restituire `"hard"`.

# --hints--

Dovresti creare una `getDifficultyLevel` funzione.

```js
assert.isFunction(getDifficultyLevel);
```

La tua `getDifficultyLevel` funzione dovrebbe avere un solo parametro.

```js
assert.lengthOf(getDifficultyLevel, 1);
```

La tua `getDifficultyLevel` funzione dovrebbe restituire `"easy"` quando il tempo di cottura è minore o uguale a `30`.

```js
assert.strictEqual(getDifficultyLevel(10), "easy");
assert.strictEqual(getDifficultyLevel(20), "easy");
assert.strictEqual(getDifficultyLevel(29), "easy");
assert.strictEqual(getDifficultyLevel(30), "easy");
```

La tua `getDifficultyLevel` funzione dovrebbe restituire `"medium"` quando il tempo di cottura è maggiore di `31` e minore o uguale a `60`.

```js
assert.strictEqual(getDifficultyLevel(31), "medium");
assert.strictEqual(getDifficultyLevel(40), "medium");
assert.strictEqual(getDifficultyLevel(50), "medium");
assert.strictEqual(getDifficultyLevel(60), "medium");
```

La tua `getDifficultyLevel` funzione dovrebbe restituire `"hard"` quando il tempo di cottura è maggiore di `60`.

```js
assert.strictEqual(getDifficultyLevel(61), "hard");
assert.strictEqual(getDifficultyLevel(65), "hard");
assert.strictEqual(getDifficultyLevel(70), "hard");
assert.strictEqual(getDifficultyLevel(75), "hard");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

--fcc-editable-region--

--fcc-editable-region--
```
