---
id: 66fbcf750a62784cf11f5632
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Crea una `getTotalIngredients` funzione che prenda un singolo argomento, rappresentante un array con ingredienti, e restituisca il numero di ingredienti dell'array passato alla funzione.

# --hints--

Dovresti creare una `getTotalIngredients` funzione.

```js
assert.isFunction(getTotalIngredients);
```

La tua `getTotalIngredients` funzione dovrebbe avere un solo parametro.

```js
assert.lengthOf(getTotalIngredients, 1);
```

La tua `getTotalIngredients` funzione dovrebbe restituire un numero.

```js
assert.isNumber(getTotalIngredients(recipe1.ingredients));
```

`getTotalIngredients(recipe1.ingredients)` dovrebbe restituire `4`.

```js
assert.strictEqual(getTotalIngredients(recipe1.ingredients), 4);
```

`getTotalIngredients(recipe2.ingredients)` dovrebbe restituire `5`.

```js
assert.strictEqual(getTotalIngredients(recipe2.ingredients), 5);
```

`getTotalIngredients(recipe3.ingredients)` dovrebbe restituire `3`.

```js
assert.strictEqual(getTotalIngredients(recipe3.ingredients), 3);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

--fcc-editable-region--

--fcc-editable-region--
```
