---
id: 66fbcdfed6848e48947dbbc9
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop, creerai un tracker di ricette usando oggetti JavaScript.

Inizia creando un array vuoto chiamato `recipes`. Questo è l'array in cui aggiungerai gli oggetti ricetta più avanti.

# --hints--

Dovresti creare un array chiamato `recipes`.

```js
assert.isArray(recipes);
```

Il tuo array `recipes` dovrebbe essere vuoto.

```js
assert.isEmpty(recipes);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
