---
id: 66fbcf750a62784cf11f562e
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Crea un oggetto `recipe2` con le seguenti proprietà e valori:

| Chiave | Valore   |
| ----------- | ------- |
| `name`       | `Chicken Curry` |
| `ingredients`       | `["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]` |
| `cookingTime`       | `42` |
| `totalIngredients`   | `null`   |
| `difficultyLevel`   | `""`   |

# --hints--

Dovresti creare un oggetto chiamato `recipe2`.

```js
assert.isObject(recipe2);
```

Il tuo oggetto `recipe2` dovrebbe avere una proprietà `name`.

```js
assert.property(recipe2, "name");
```

La tua proprietà `name` dovrebbe essere impostata su `Chicken Curry`.

```js
assert.equal(recipe2.name, "Chicken Curry");
```

Il tuo oggetto `recipe2` dovrebbe avere una proprietà `ingredients`.

```js
assert.property(recipe2, "ingredients");
```

La tua proprietà `ingredients` dovrebbe avere come valore un array.

```js
assert.isArray(recipe2.ingredients);
```

L'array valore della tua proprietà `ingredients` dovrebbe contenere `chicken breast`, `coconut milk`, `curry powder`, `onion` e `garlic`.

```js
assert.deepEqual(recipe2.ingredients, ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]);
```

Il tuo oggetto `recipe2` dovrebbe avere una proprietà `cookingTime`.

```js
assert.property(recipe2, "cookingTime");
```

Il valore della tua proprietà `cookingTime` dovrebbe essere un numero.

```js
assert.isNumber(recipe2.cookingTime);
```

La tua proprietà `cookingTime` dovrebbe essere impostata su `42`.

```js
assert.equal(recipe2.cookingTime, 42);
```

Il tuo oggetto `recipe2` dovrebbe avere una proprietà `totalIngredients`.

```js
assert.property(recipe2, "totalIngredients");
```

La tua proprietà `totalIngredients` dovrebbe essere impostata su `null`.

```js
assert.isNull(recipe2.totalIngredients);
```

Il tuo oggetto `recipe2` dovrebbe avere una proprietà `difficultyLevel`.

```js
assert.property(recipe2, "difficultyLevel");
```

La tua proprietà `difficultyLevel` dovrebbe essere impostata su una stringa vuota.

```js
assert.deepEqual(recipe2.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
  };

--fcc-editable-region--

--fcc-editable-region--
```
