---
id: 66ff02792b3042a2b811f475
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Aggiungi le seguenti proprietà all'oggetto `recipe1`:

| Chiave | Valore   |
| ----------- | ------- |
|`cookingTime`|`22`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

Le proprietà con `null` e valori stringa vuota verranno aggiornate più avanti dopo che le avrai calcolate.

# --hints--

Il tuo oggetto `recipe1` dovrebbe avere una proprietà `cookingTime`.

```js
assert.property(recipe1, "cookingTime");
```

La tua proprietà `cookingTime` dovrebbe essere impostata su `22`.

```js
assert.equal(recipe1.cookingTime, 22);
```

Il tuo oggetto `recipe1` dovrebbe avere una proprietà `totalIngredients`.

```js
assert.property(recipe1, "totalIngredients");
```

La tua proprietà `totalIngredients` dovrebbe essere impostata su `null`.

```js
assert.isNull(recipe1.totalIngredients);
```

Il tuo oggetto `recipe1` dovrebbe avere una proprietà `difficultyLevel`.

```js
assert.property(recipe1, "difficultyLevel");
```

La tua proprietà `difficultyLevel` dovrebbe essere impostata su una stringa vuota.

```js
assert.deepEqual(recipe1.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
--fcc-editable-region--
  
--fcc-editable-region--
};
```
