---
id: 670e4f45f7116c0f216a5177
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Prima di andare avanti, dovresti esercitarti a consultare le proprietà di un oggetto.

Puoi usare sia la notazione a punto (`.`) sia quella a parentesi (`[]`) per farlo. Ecco un esempio:

```js
const person = {
  name: "John",
  age: 30,
  job: "Software Engineer"
};

console.log(person.name); // John
console.log(person['age']);  // 30
```

Consulta la proprietà `name` di `recipe1` e assegnala alla variabile `recipe1Name`.

Poi, consulta la proprietà `cookingTime` di `recipe2` e assegnala alla variabile `recipe2CookingTime`.

Infine, consulta la proprietà `ingredients` di `recipe3` e assegnala alla variabile `recipe3Ingredients`.

Assicurati che tutte le variabili create vengano stampate nella console.

# --hints--

Dovresti creare una variabile `recipe1Name`.

```js
assert.isDefined(recipe1Name);
```

Dovresti assegnare il valore della proprietà `name` di `recipe1` alla tua variabile `recipe1name`.

```js
assert.strictEqual(recipe1Name, recipe1.name);
```

Dovresti stampare `recipe1Name` nella console.

```js
assert.match(code, /console\.log\(\s*recipe1Name\s*\)/);
```

Dovresti creare una variabile `recipe2CookingTime`.

```js
assert.isDefined(recipe2CookingTime);
```

Dovresti assegnare il valore della proprietà `cookingTime` di `recipe2` alla tua variabile `recipe2CookingTime`.

```js
assert.strictEqual(recipe2CookingTime, recipe2.cookingTime);
```

Dovresti stampare `recipe2CookingTime` nella console.

```js
assert.match(code, /console\.log\(recipe2CookingTime\)/);
```

Dovresti creare una variabile `recipe3Ingredients`.

```js
assert.isDefined(recipe3Ingredients);
```

Dovresti assegnare il valore della proprietà `ingredients` di `recipe3` alla tua variabile `recipe3Ingredients`.

```js
assert.strictEqual(recipe3Ingredients, recipe3.ingredients);
```

Dovresti stampare `recipe3Ingredients` nella console.

```js
assert.match(code, /console\.log\(recipe3Ingredients\)/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
