---
id: 69161f4c0c9c8c48b3c843c5
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Crea un oggetto `recipe3` con le seguenti proprietà e valori:

|Chiave|Valore|
|-----------|-------|
|`name`|`Vegetable Stir Fry`|
|`ingredients`|`["broccoli", "carrot", "bell pepper"]`|
|`cookingTime`|`15`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

# --hints--

Dovresti creare un oggetto chiamato `recipe3`.

```js
assert.isObject(recipe3);
```

Il tuo oggetto `recipe3` dovrebbe avere una proprietà `name`.

```js
assert.property(recipe3, "name");
```

La tua proprietà `name` dovrebbe essere impostata su `Vegetable Stir Fry`.

```js
assert.equal(recipe3.name, "Vegetable Stir Fry");
```

Il tuo oggetto `recipe3` dovrebbe avere una proprietà `ingredients`.

```js
assert.property(recipe3, "ingredients");
```

La tua proprietà `ingredients` dovrebbe avere un valore di tipo array.

```js
assert.isArray(recipe3.ingredients);
```

Il valore array della tua proprietà `ingredients` dovrebbe contenere `broccoli`, `carrot` e `bell pepper`.

```js
assert.deepEqual(recipe3.ingredients, ["broccoli", "carrot", "bell pepper"]);
```

Il tuo oggetto `recipe3` dovrebbe avere una proprietà `cookingTime`.

```js
assert.property(recipe3, "cookingTime");
```

Il valore della tua proprietà `cookingTime` dovrebbe essere un numero.

```js
assert.isNumber(recipe3.cookingTime);
```

La tua proprietà `cookingTime` dovrebbe essere impostata su `15`.

```js
assert.equal(recipe3.cookingTime, 15);
```

Il tuo oggetto `recipe3` dovrebbe avere una proprietà `totalIngredients`.

```js
assert.property(recipe3, "totalIngredients");
```

La tua proprietà `totalIngredients` dovrebbe essere impostata su `null`.

```js
assert.isNull(recipe3.totalIngredients);
```

Il tuo oggetto `recipe3` dovrebbe avere una proprietà `difficultyLevel`.

```js
assert.property(recipe3, "difficultyLevel");
```

La tua proprietà `difficultyLevel` dovrebbe essere impostata su una stringa vuota.

```js
assert.deepEqual(recipe3.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
