---
id: 69bac027df64d4433c708d44
title: 第 12 步
challengeType: 1
dashedName: step-12
---

# --description--

在重新路由溢出时，你应该将参与者发送到 `gates` 数组中的下一个闸门。为此，你需要查找下一个闸门的索引。通常，你可以将当前索引加 1。但是，如果当前闸门是数组中的最后一个，你需要回绕到第一个闸门。

你可以使用取模操作符（`%`）来处理这个问题：

```js
const nextIndex = (currentIndex + 1) % array.length;
```

这是因为当 `currentIndex + 1` 等于 `array.length` 时，结果变为 0。

使用这种方法，创建一个名为 `nextGateIndex` 的变量，用于保存 `gates` 数组中下一个闸门的索引。

# --hints--

你应该使用取模技巧，从 `currentIndex` 中访问并赋值 `gates` 数组的下一个索引给名为 `nextGateIndex` 的变量。遵循示例中的同一行模式。

```js
const rawFuncStr = __helpers.removeJSComments(rerouteOverflow.toString());
assert.match(
  rawFuncStr,
  /(var|const|let)\s+nextGateIndex/,
  "You must have at least one space between the declaration keyword (let/const) and 'nextGateIndex'."
);

const cleaned = __helpers.removeWhiteSpace(rawFuncStr);
const solution = /(var|const|let)nextGateIndex=\((currentIndex\+1|1\+currentIndex)\)%gates\.length;?}$/;

assert(
  solution.test(cleaned),
  "Inside your `rerouteOverflow` function, use the modulo technique to create a variable named `nextGateIndex` that stores the index of the next gate."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
--fcc-editable-region--

--fcc-editable-region--
}
```
