---
id: 69babdf3f896711c11516373
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

在每个模拟时钟周期中：

- 一定数量的参与者会从其 `queue` 到达门口。
- 门口根据其 `capacity` 处理参与者。
- 如果到达的参与者超过门口能处理的数量，一些参与者将会留下（溢出）。

你现在将创建一个函数，用于处理单个门口在单个时钟周期的这套逻辑。

创建一个名为 `processGateFlow` 的空函数，该函数接受两个参数：

- `gate`：正在处理的门口对象。
- `tickIndex`：`queue` 数组中当前模拟时钟周期的位置。

# --hints--

你应该创建一个名为 `processGateFlow` 的函数。

```js
assert.isFunction(processGateFlow);
```

你的函数 `processGateFlow` 应该有参数 `gate` 和 `tickIndex`。

```js
const regex = __helpers.functionRegex('processGateFlow', ['gate', 'tickIndex']);
assert.match(__helpers.removeJSComments(processGateFlow.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

--fcc-editable-region--

--fcc-editable-region--
```
