---
id: 69baafd6d683dd9921b75db3
title: 第一步
challengeType: 1
dashedName: step-1
---

# --description--

在本次工作坊中，你将使用 JavaScript 来模拟音乐节参与者的流动。

你将处理两个描述音乐节大门在早班和晚班期间情况的数据集：

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];
```

每个大门对象包含以下属性：

- `id`：一个标识大门的字串。
- `capacity`：大门每个时刻能处理的参与者数量。
- `queue`：一个数字数组，表示在特定时刻到达大门的参与者数量。

首先，创建一个名为 `initializeThroughput` 的空函数，该函数接受一个名为 `gates` 的参数。

# --hints--

你应该创建一个名为 `initializeThroughput` 的函数。

```js
assert.isFunction(initializeThroughput);
```

你的 `initializeThroughput` 函数应该有一个 `gates` 参数。

```js
const regex = __helpers.functionRegex('initializeThroughput', ['gates']);
assert.match(__helpers.removeJSComments(initializeThroughput.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

--fcc-editable-region--

--fcc-editable-region--
```
