---
id: 69bac067cd15c4685d07906d
title: 第 15 步
challengeType: 1
dashedName: step-15
---

# --description--

既然你已经有了处理单个时钟周期和处理溢出的函数，现在是时候为单个闸门创建一个控制函数了。这个函数将在模拟的单个时钟周期内处理一个闸门的所有操作。

创建一个名为 `handleGateAtTick` 的空函数，带有以下参数：

- `gates`：完整的闸门对象数组。
- `gate`：当前正在处理的闸门。
- `tickIndex`：模拟的当前时钟周期索引。
- `throughputSummary`：一个跟踪每个闸门处理的总参与者的对象。

# --hints--

你应该创建一个名为 `handleGateAtTick` 的函数。

```js
assert.isFunction(handleGateAtTick);
```

你的 `handleGateAtTick` 函数应包含参数 `gates`、`gate`、`tickIndex` 和 `throughputSummary`。

```js
const regex = __helpers.functionRegex('handleGateAtTick', ['gates', 'gate', 'tickIndex', 'throughputSummary']);
assert.match(__helpers.removeJSComments(handleGateAtTick.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

--fcc-editable-region--

--fcc-editable-region--
```
