---
id: 69babfd65c4c4632a0c76f63
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

当一个闸口在一个时钟周期内无法处理所有参与者时，一些参与者会留在队列中。为了解决这种溢出情况，你将创建一个函数，将溢出的参与者重新引导到另一个闸口。

创建一个名为 `rerouteOverflow` 的空函数，带有以下参数：

- `gates`：包含所有闸口对象的完整数组。
- `currentGate`：发生溢出的闸口。
- `tickIndex`：队列中当前的时钟周期位置。
- `overflowAmount`：无法处理的参与者数量。

# --hints--

你应该创建一个名为 `rerouteOverflow` 的函数。

```js
assert.isFunction(rerouteOverflow);
```

你的 `rerouteOverflow` 函数应包含参数 `gates`、`currentGate`、`tickIndex` 和 `overflowAmount`。

```js
const regex = __helpers.functionRegex('rerouteOverflow', ['gates', 'currentGate', 'tickIndex', 'overflowAmount']);
assert.match(__helpers.removeJSComments(rerouteOverflow.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

--fcc-editable-region--

--fcc-editable-region--
```
