---
id: 69bab8292750d84cac5b52a0
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

在本次工作坊中，吞吐量指的是在单次模拟（早晨或晚上）中，所有时刻内每个闸门处理的总参与者数量。你的 `initializeThroughput` 函数将准备一个对象，用于跟踪每个闸门处理了多少参与者。

首先，创建一个名为 `summary` 的空对象。

# --hints--

你的 `initializeThroughput` 函数应该创建一个空对象，并赋值给名为 `summary` 的变量。

```js
const explorer = await __helpers.Explorer(code);
const { initializeThroughput } = explorer.allFunctions;
const { summary } = initializeThroughput?.variables ?? {};
const validExpressions = ["{}", "new Object()"];
assert.isTrue(
  validExpressions.some((expression) => summary?.value.matches(expression))
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
