---
id: 69babea99725b43a01adb6d1
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

首先，你需要获取当前时刻到达的参与者人数。访问输入门的 `queue` 数组中索引为 `tickIndex` 的值，并将其赋值给名为 `currentTickQueue` 的变量。

# --hints--

你应该将 `gate.queue[tickIndex]` 处的元素赋值给名为 `currentTickQueue` 的变量。

```js
const explorer = await __helpers.Explorer(code);
const { processGateFlow } = explorer.allFunctions;
const { currentTickQueue } = processGateFlow?.variables ?? {};
assert.isTrue(currentTickQueue?.value.matches("gate.queue[tickIndex]"));
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
