---
id: 69bac0d5a0b16122a2a938a4
title: 第 19 步
challengeType: 1
dashedName: step-19
---

# --description--

接下来，更新输入 `throughputSummary` 对象，以反映当前时刻该闸口成功处理的参与者数量。具体来说，取 `processed` 中 `result` 的值，并将其添加到 `throughputSummary` 中索引为 `gate.id` 的元素上。

这一步确保吞吐量汇总会在多个时刻累积参与者数量。

# --hints--

你应该将 `result.processed` 添加到 `throughputSummary[gate.id]`。

```js
const testGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];
const throughputSummary = { North: 0, East: 0, South: 0, West: 0 };
const console_log = console.log;
console.log = () => {};

try {
  const northGate = testGates[0];
  handleGateAtTick(testGates, northGate, 0, throughputSummary);
  assert.strictEqual(
    throughputSummary["North"],
    3,
    "North gate should have 3 attendees processed at tick 0 (capacity 5, queue 3)"
  );
} finally {
  console.log = console_log;
}
```

你的 `handleGateAtTick` 函数应该在多个时刻累积处理过的参与者。

```js
const testGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];

const throughputSummary = { North: 0, East: 0, South: 0, West: 0 };
const console_log = console.log;
console.log = () => {};

try {
  const northGate = testGates[0];
  handleGateAtTick(testGates, northGate, 0, throughputSummary);
  handleGateAtTick(testGates, northGate, 1, throughputSummary);
  assert.strictEqual(
    throughputSummary["North"],
    8,
    "North gate should accumulate processed attendees: 3 from tick 0 + 5 from tick 1"
  );
} finally {
  console.log = console_log;
}
```

你的 `handleGateAtTick` 函数应该只更新正在处理的特定闸口的吞吐量。

```js
const testGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];

const throughputSummary = { North: 0, East: 0, South: 0, West: 0 };
const console_log = console.log;
console.log = () => {};
try {
  const westGate = testGates[3];
  handleGateAtTick(testGates, westGate, 0, throughputSummary);
  assert.strictEqual(
    throughputSummary["West"],
    2,
    "West gate should have 2 attendees processed (capacity 2, queue 4)"
  );
  assert.strictEqual(
    throughputSummary["North"],
    0,
    "North gate should remain unmodified"
  );
  assert.strictEqual(
    throughputSummary["East"],
    0,
    "East gate should remain unmodified"
  );
  assert.strictEqual(
    throughputSummary["South"],
    0,
    "South gate should remain unmodified"
  );
} finally {
  console.log = console_log;
}
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

function handleGateAtTick(gates, gate, tickIndex, throughputSummary) {
  console.log("\nProcessing " + gate.id + "...");
  console.log(
    gate.queue[tickIndex] + " attendees arriving."
  );
  const result = processGateFlow(gate, tickIndex);
--fcc-editable-region--

--fcc-editable-region--
}
```
