---
id: 69d08c139139f65a17cc615c
title: 第 30 步
challengeType: 1
dashedName: step-30
---

# --description--

在你的 `while` 循环内：

- 记录消息 `"\nTick " + (tickIndex + 1)`。
- 遍历 `gates` 数组中的每个 gate，并使用以下参数调用 `handleGateAtTick()`：
  - `gates`：完整的 gates 输入数组。
  - `gate`：当前的 gate 对象。
  - `tickIndex`：当前的 tick。
  - `throughputSummary`：跟踪每个 gate 处理的总参与者数的局部对象。
- 将 `tickIndex` 增加 1，以进入下一个 tick。

# --hints--

你应该将以下消息 `"\nTick " + (tickIndex + 1)` 记录到控制台。

```js
const testGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];

const spy = __helpers.spyOn(console, "log");

try {
  simulateFestival(testGates, "Morning");
  const tickMessages = spy.calls.filter(call => call[0].includes("Tick"));
  assert.lengthOf(
    tickMessages,
    4,
    "Should log a tick message for each tick (4 ticks total)"
  );
  assert.match(
    tickMessages[0][0],
    /\nTick 1/,
    "First tick should log '\\nTick 1'"
  );
  assert.match(
    tickMessages[1][0],
    /\nTick 2/,
    "Second tick should log '\\nTick 2'"
  );
  assert.match(
    tickMessages[2][0],
    /\nTick 3/,
    "Third tick should log '\\nTick 3'"
  );
  assert.match(
    tickMessages[3][0],
    /\nTick 4/,
    "Fourth tick should log '\\nTick 4'"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

在记录到控制台后，你应该遍历 `gates` 数组中的每个 gate，并用 `handleGateAtTick()`、`gates`、`gate` 和 `tickIndex` 调用 `throughputSummary`。

```js
const testGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];

const originalHandleGateAtTick = handleGateAtTick;
let handleGateCallCount = 0;

try {
  handleGateAtTick = function (...args) {
    handleGateCallCount++;
    return originalHandleGateAtTick(...args);
  };

  const consoleSpy = __helpers.spyOn(console, "log");
  
  try {
    simulateFestival(testGates, "Morning");
    
    const expectedCallCount = testGates.length * testGates[0].queue.length;
    
    assert.strictEqual(
      handleGateCallCount,
      expectedCallCount,
      `handleGateAtTick should be called ${expectedCallCount} times (4 gates × 4 ticks)`
    );

  } finally {
    consoleSpy.restore();
  }

} catch (err) {
  assert.fail(err);
} finally {
  handleGateAtTick = originalHandleGateAtTick;
}
```

在遍历 `gates` 数组后，你应该将 `tickIndex` 增加 1。

```js
const testGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];

const consoleSpy = __helpers.spyOn(console, "log");

try {
  simulateFestival(testGates, "Morning");
  
  const allLogs = consoleSpy.calls.length;
  
  assert.isTrue(
    allLogs > 0,
    "Function should complete and log messages (without infinite loop)"
  );
  
  const tickMessages = consoleSpy.calls.filter(call => call[0].includes("Tick"));
  assert.lengthOf(
    tickMessages,
    4,
    "tickIndex should be incremented properly to complete all 4 ticks without infinite loop"
  );

} catch (err) {
  assert.fail(err);
} finally {
  consoleSpy.restore();
}
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

function handleGateAtTick(gates, gate, tickIndex, throughputSummary) {
  console.log("\nProcessing " + gate.id + "...");
  console.log(
    gate.queue[tickIndex] + " attendees arriving."
  );
  const result = processGateFlow(gate, tickIndex);
  throughputSummary[gate.id] += result.processed;
  if (result.overflow > 0) {
    console.log(
      "Overflow of " + result.overflow +
      " attendees. Rerouting..."
    );
    rerouteOverflow(gates, gate, tickIndex, result.overflow);
  }
}

function printSummary(summary) {
  console.log("\nThroughput Summary");
  for (const gateId in summary) {
    console.log(
      gateId + ": " + summary[gateId] +
      " attendees processed"
    );
  }
}

function simulateFestival(gates, timeBlock) {
  console.log("\n" + timeBlock + " Simulation");
  const throughputSummary = initializeThroughput(gates);
  const maxTicks = gates[0].queue.length;
  let tickIndex = 0;
  while (tickIndex < maxTicks) {
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
