---
id: 5dc17d3bf86c76b9248c6eb4
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Unda kipengele cha `p` chini ya kipengele chako cha `h2` na kimpatie maandishi yafuatayo:

`Everyone loves cute cats online!`

# --hints--

Kipengele chako cha `p` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia ifuatayo: `<elementName>`.

```js
assert.exists(document.querySelector('p'));
```

Lebos zako za `p` zinapaswa kuwa kwa herufi ndogo. Kwa mkataba, lebo zote za HTML huandikwa kwa herufi ndogo.

```js
assert.notMatch(code, /<\/?P>/);
```

Kipengele chako cha `p` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert.match(code, /<\/p\>/);
```

Maandishi ya kipengele chako cha `p` yanapaswa kuwa `Everyone loves cute cats online!` Umeacha maandishi au kuna makosa ya tahajia.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  ?.innerText?.trim().replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

Kipengele chako cha `p` kinapaswa kuwa chini ya kipengele cha `h2`. Umeviweka kwa mpangilio usio sahihi.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H2'), collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
