---
id: 5dc1798ff86c76b9248c6eb3
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Chini ya kipengele cha `h1`, ongeza kipengele cha `h2` chenye maandishi haya:

`Cat Photos`

# --hints--

Kipengele chako cha `h1` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia hii: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Kipengele chako cha `h1` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina sintaksia hii: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Unapaswa kuwa na kipengele kimoja tu cha `h1`. Ondoa ziada.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Maandishi ya kipengele chako cha `h1` yanapaswa kuwa 'CatPhotoApp'. Umeacha maandishi au kuna makosa ya tahajia.

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```

Lebos za `h2` zinapaswa kuwa kwa herufi ndogo. Kwa mkataba, lebo zote za HTML huandikwa kwa herufi ndogo.

```js
assert.notMatch(code, /<\/?H2>/);
```

Kipengele chako cha `h2` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia hii: `<elementName>`.

```js
assert.exists(document.querySelector('h2'));
```

Kipengele chako cha `h2` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara baada ya herufi ya `<`.

```js
assert.match(code, /<\/h2\>/);
```

Maandishi ya kipengele chako cha `h2` yanapaswa kuwa `Cat Photos`. Weka maandishi `Cat Photos` kati ya lebo za kufungua na kufunga za `h2` pekee.

```js
assert.equal(document.querySelector('h2')?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat photos');
```

Kipengele chako cha `h2` kinapaswa kuwa chini ya kipengele cha `h1`. Kipengele cha `h1` kina umuhimu zaidi na kinapaswa kuwa juu ya kipengele cha `h2`.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H1'), collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
