---
id: 5dc174fcf86c76b9248c6eb2
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utaendelea kufanya kazi na vipengele vya msingi vya HTML kama vichwa vya sehemu, aya, na orodha kwa kujenga programu ya picha ya paka.

Anza warsha kwa kuongeza kipengele cha `h1` chenye maandishi ya `CatPhotoApp`.

# --hints--

Maneno `CatPhotoApp` yanapaswa kuwepo katika msimbo. Labda unataka kukagua tahajia yako.

```js
assert.match(code, /catphotoapp/i);
```

Kipengele chako cha `h1` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia hii: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Lebos za `h1` zinapaswa kuwa kwa herufi ndogo. Kwa mkataba, lebo zote za HTML huandikwa kwa herufi ndogo.

```js
assert.notMatch(code, /<\/?H1>/);
```

Kipengele chako cha `h1` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina sintaksia hii: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Maneno ya kipengele cha `h1` yanapaswa kuwa `CatPhotoApp`. Huenda umeacha maneno hayo, kuna makosa ya tahajia, au hayako kati ya lebo ya kufungua na lebo ya kufunga ya kipengele cha `h1`.

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
