---
id: 5dc24614f86c76b9248c6ebd
title: Hatua 10
challengeType: 0
dashedName: step-10
---

# --description--

Unaweza kuunganisha na ukurasa mwingine kwa kutumia kipengele cha nanga (`a`). 

Hapa kuna mfano wa kuunganisha na `https://www.freecodecamp.org`:

```html
<a href="https://www.freecodecamp.org"></a>
```

Ongeza kipengele cha nanga baada ya aya kinachounganisha na `https://freecatphotoapp.com`. Hadi sasa, kiungo hakitaonekana kwenye mtazamo wa awali.

# --hints--

Kipengele chako cha nanga (`a`) kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia hii: `<elementName>`.

```js
assert.exists(document.querySelector('a'));
```

Kipengele chako cha nanga (`a`) kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert.match(code, /<\/a\>/);
```

Kipengele chako cha nanga (`a`) kinapaswa kuwa chini ya kipengele cha `p`. Umeviweka kwa mpangilio usio sahihi.

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('A'));
```

Kipengele chako cha nanga (`a`) hakina sifa ya `href`. Hakikisha kuna nafasi baada ya jina la lebo ya kufungua na/au kuna nafasi kabla ya majina yote ya sifa.

```js
assert.isTrue(document.querySelector('a')?.hasAttribute('href'));
```

Kipengele chako cha nanga (`a`) kinapaswa kuunganisha na `https://freecatphotoapp.com`. Umeacha URL au umefanya makosa ya tahajia.

```js
assert.equal(
  document.querySelector('a')?.getAttribute('href'),
    'https://freecatphotoapp.com'
);
```

Ingawa umeweka sifa ya `a` ya kipengele cha nanga (`href`) kwa kiungo sahihi, inashauriwa kila mara kuzunguka thamani ya sifa na alama za nukuu.

```js
assert.notMatch(
  code, /\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
