---
id: 5dc24073f86c76b9248c6ebb
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Sifa za kipengele za HTML ni maneno maalum yanayotumika ndani ya lebo ya kufungua ya kipengele kudhibiti tabia ya kipengele hicho. Sifa ya kipengele `src` katika kipengele cha `img` inaelezea URL ya picha (mahali ambapo picha iko).

Hapa kuna mfano wa kipengele cha `img` chenye sifa ya kipengele `src` inayorejelea nembo ya freeCodeCamp:

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

Ndani ya kipengele cha `img` kilichopo, ongeza sifa ya kipengele `src` yenye URL hii:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `img`. Huenda umeondoa kipengele cha `img` au hujazunguka thamani ya sifa ya kipengele `src` kwa nukuu.

```js
assert.exists(document.querySelector('img'));
```

Kipengele chako cha `img` kinapaswa kuwa na sifa ya kipengele `src`. Huenda umeacha sifa hiyo au kuna makosa ya tahajia. Hakikisha kuna nafasi kati ya jina la kipengele na jina la sifa ya kipengele.

```js
assert.exists(document.querySelector('img')?.src);
```

Sifa ya kipengele `img` ya kipengele cha `src` inapaswa kuwekwa kuwa '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'. Huenda umeacha URL hiyo au kuna makosa ya tahajia. Herufi kubwa na ndogo za URL ni muhimu.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

Ingawa umeweka sifa ya kipengele `img` ya kipengele cha `src` kwa URL sahihi, inashauriwa kila mara kuzunguka thamani ya sifa ya kipengele kwa alama za nukuu.

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

Kipengele chako cha `img` kinapaswa kufungwa kwa `>` au `/>`.

```js
assert.match(code, /<img\s+src\s*=\s*("|')https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```
