---
id: 5dc23f9bf86c76b9248c6eba
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Unaweza kuongeza picha kwenye tovuti yako kwa kutumia kipengele cha `img`. Vipengele vya `img` vina lebo ya kufungua bila lebo ya kufunga. Kipengele kisicho na lebo ya kufunga kinajulikana kama <dfn>kipengele tupu</dfn>.

Ongeza kipengele cha `img` chini ya kipengele cha `p`. Hadi sasa, picha haitajitokeza kwenye kivinjari.

# --hints--

Kipengele chako cha `img` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia hii: `<elementName>`.

```js
assert.exists(document.querySelector('img'));
```

Kipengele chako cha `img` hakipaswi kuwa na lebo ya kufunga `</img>`.

```js
assert.notMatch(code, /<\/img\>/);
```

Unapaswa kuwa na kipengele kimoja tu cha `img`. Ondoa ziada zozote.

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

Kipengele chako cha `img` kinapaswa kuwa chini ya kipengele cha `p`. Umeviweka kwa mpangilio usio sahihi.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
