---
id: 5dfb6a35eacea3f48c6300b4
title: Step 26
challengeType: 0
dashedName: step-26
---

# --description--

Kipengele cha maelezo ya picha (`figcaption`) hutumika kuongeza maelezo ya picha kuelezea picha iliyomo ndani ya kipengele cha `figure`.

Hapa kuna mfano wa kipengele cha `figcaption` chenye maelezo ya `A cute cat`:

```html
<figure>
  <img src="image.jpg" alt="A description of the image">
  <figcaption>A cute cat</figcaption>
</figure>
```

Baada ya picha iliyopangwa ndani ya kipengele cha `figure`, ongeza kipengele cha `figcaption` chenye maandishi yaliyowekwa kuwa:

`Cats love lasagna.`

# --hints--

Kipengele chako cha `figcaption` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia ifuatayo: `<elementName>`.

```js
assert.exists(document.querySelector('figcaption'));
```

Kipengele chako cha `figcaption` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert.match(code, /<\/figcaption\>/);
```

Kipengele cha `figcaption` kinapaswa kupangwa ndani ya kipengele cha `figure`.

```js
assert.exists(document.querySelector('figure > figcaption'));
```

Kipengele cha lasagna cha `img` kinapaswa kupangwa ndani ya kipengele cha `figure`.

```js
assert.exists(document.querySelector('figure > img'))
assert.equal(
    document.querySelector('figure > img')?.getAttribute('src').toLowerCase(),
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

Kipengele cha `figcaption` kilichopangwa ndani ya kipengele cha `figure` kinapaswa kuwa chini ya kipengele cha `img`. Umeviweka kwa mpangilio usio sahihi.

```js
assert.equal(
  document.querySelector('figcaption')?.previousElementSibling.nodeName, 'IMG'
);
```

Maandishi ya kipengele chako cha `figcaption` yanapaswa kuwa `Cats love lasagna.`. Umeacha maandishi au kuna makosa ya tahajia.

```js
assert.match(
  document.querySelector('figcaption')?.innerText?.trim(), /Cats love lasagna.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          
--fcc-editable-region--          
        </figure>
      </section>
    </main>
  </body>
</html>
```
