---
id: 5dfb6250eacea3f48c6300b2
title: Step 24
challengeType: 0
dashedName: step-24
---

# --description--

Baada ya orodha isiyopangwa, ongeza picha mpya yenye thamani ya sifa ya kipengele `src` kuwekwa kuwa:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

Na thamani ya sifa ya kipengele `alt` kuwa:

`A slice of lasagna on a plate.`

# --hints--

Kuwa na kipengele cha `img` mara moja baada ya lebo ya kufunga `</ul>`.

```js
assert.equal(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.nodeName, 'IMG');
```

Picha mpya haina sifa ya kipengele `alt`. Hakikisha kuna nafasi baada ya jina la lebo ya kufungua na/au kuna nafasi kabla ya majina yote ya sifa za kipengele.

```js
assert.isTrue(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.hasAttribute('alt'));
```

Picha mpya inapaswa kuwa na thamani ya `alt` ya `A slice of lasagna on a plate.`. Hakikisha thamani ya sifa ya kipengele `alt` imezungukwa na alama za nukuu.

```js
assert.match(
  document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.getAttribute('alt')?.replace(/\s+/g, ' '), /^A slice of lasagna on a plate\.?$/i 
);
```

Picha mpya haina sifa ya kipengele `src`. Hakikisha kuna nafasi baada ya jina la lebo ya kufungua na/au kuna nafasi kabla ya majina yote ya sifa za kipengele.

```js
assert.isTrue(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.hasAttribute('src'));
```

Picha mpya inapaswa kuwa na thamani ya `src` ya `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`. Hakikisha thamani ya sifa ya kipengele `src` imezungukwa na alama za nukuu.

```js
assert.equal(
  document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.getAttribute('src'),
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

Ingawa umeweka sifa ya kipengele cha `src` ya picha mpya kuwa URL sahihi, inashauriwa kila mara kuzunguka thamani ya sifa ya kipengele na alama za nukuu.

```js
assert.notMatch(code, /\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
