---
id: 5dfb655eeacea3f48c6300b3
title: Step 25
challengeType: 0
dashedName: step-25
---

# --description--

Kipengele cha `figure` kinakilisha maudhui yaliyojitegemea na kitakuwezesha kuhusisha picha na maelezo ya picha.

Panga ndani picha uliyoiweka hivi karibuni ndani ya kipengele cha `figure`.

# --hints--

Kipengele chako cha `figure` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia ifuatayo: `<elementName>`.

```js
assert.exists(document.querySelector('figure'));
```

Kipengele chako cha `figure` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert.match(code, /<\/figure\>/);
```

Kipengele cha `figure` kinapaswa kuwepo juu kabisa ya lebo ya kufunga ya kipengele cha `section` cha pili.

```js
assert.equal(document.querySelectorAll('section')[1]?.lastElementChild.nodeName, 'FIGURE');
```

Kipengele cha lasagna cha `img` kinapaswa kupangwa ndani ya kipengele cha `figure`.

```js
assert.exists(document.querySelector('figure > img'));
assert.equal(
    document.querySelector('figure > img').getAttribute('src').toLowerCase(),
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        
        <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
        
--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```
