---
id: 5ef9b03c81a63668521804e8
title: Step 36
challengeType: 0
dashedName: step-36
---

# --description--

Panga ndani kipengele cha `p` chenye maandishi ya `No Copyright - freeCodeCamp.org` ndani ya kipengele cha `footer`.

# --hints--

Umefuta kipengele cha `footer` au kinakosa lebo ya kufungua au lebo ya kufunga.

```js
assert.exists(document.querySelector('footer'));
assert.match(code, /<\/footer>/);
```

Kipengele chako cha `footer` kinapaswa kuwa na kipengele cha `p`. Hakikisha unaongeza lebo ya kufungua na lebo ya kufunga kwa kipengele cha `p`.

```js
assert.exists(document.querySelector('footer > p'));
```

Kipengele chako cha `footer` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
const pElemClosingTags = code.match(/<\/p\>/g);
assert.exists(pElemClosingTags);
assert.lengthOf(pElemClosingTags, 3);
```

Maandishi ya kipengele chako cha `p` yanapaswa kuwa `No Copyright - freeCodeCamp.org`. Ama umeacha maandishi hayo, kuna makosa ya tahajia, au hayako kati ya lebo ya kufungua na lebo ya kufunga ya kipengele cha `p`.

```js
const extraSpacesRemoved = document.querySelectorAll('footer > p')[0]?.innerText?.trim().replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /^No Copyright - freeCodeCamp\.org$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      
--fcc-editable-region--      
    </footer>
  </body>
</html>
```
