---
id: 5ef9b03c81a63668521804ec
title: Step 40
challengeType: 0
dashedName: step-40
---

# --description--

Tambua kwamba maudhui yote ya ukurasa yamepangwa ndani ya kipengele cha `html`. Kipengele cha `html` ndicho kipengele mzizi cha ukurasa wa HTML na kinazunguka maudhui yote ya ukurasa.

Pia unaweza kubainisha lugha ya ukurasa wako kwa kuongeza sifa ya kipengele ya `lang` kwenye kipengele cha `html`.

Ongeza sifa ya kipengele ya `lang` yenye thamani `en` kwenye lebo ya kufungua ya `html` ili kubainisha kwamba lugha ya ukurasa ni Kiingereza.

# --hints--

Umefuta kipengele cha `html` au kinakosa lebo ya kufungua au lebo ya kufunga.

```js
assert.match(code, /\<html.*?\>/);
assert.match(code, /\<\/html\>/);
```

Kipengele chako cha `html` kinapaswa kuwa na sifa ya kipengele ya `lang` yenye thamani `en`. Huenda umesahau sifa/thamani hiyo, au kuna makosa ya tahajia.

```js
const extraSpacesRemoved = code.replace(/\s+/g, " ");
assert.match(extraSpacesRemoved, /\<html\s+lang\s*\=\s*("?|'?)en\1\s*\>/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<html>
--fcc-editable-region--
  <head>
    <title>CatPhotoApp</title>
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
