---
id: 5ef9b03c81a63668521804ea
title: Step 38
challengeType: 0
dashedName: step-38
---

# --description--

Tambua kwamba kila kitu ulichokiweka kwenye ukurasa hadi sasa kiko ndani ya kipengele cha `body`. Vipengele vyote vya maudhui ya ukurasa vinavyopaswa kuonyeshwa kwenye ukurasa vinaenda ndani ya kipengele cha `body`. Hata hivyo, taarifa nyingine muhimu zinaenda ndani ya kipengele cha `head`.

Kipengele cha `head` kinatumika kuhifadhi metadata kuhusu hati, kama vile kichwa chake, viungo vya majalada ya mitindo, na skiripti. Metadata ni taarifa kuhusu ukurasa ambayo haionyeshwi moja kwa moja kwenye ukurasa.

Ongeza kipengele cha `head` juu ya kipengele cha `body`.

# --hints--

Umefuta kipengele cha `body` au kinakosa lebo ya kufungua au lebo ya kufunga.

```js
assert.exists(document.querySelector('body'));
assert.match(code, /<\/body>/);
```

Kipengele chako cha `head` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia ifuatayo: `<elementName>`.

```js
assert.match(code, /\<head\>/);
```

Kipengele chako cha `head` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert.match(code, /\<\/head\>/);
```

Kipengele chako cha `head` kinapaswa kuwa juu ya lebo ya kufungua ya kipengele cha `body`. Umeweka mahali pengine.

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<\/head\>\<body\>/);
```

# --seed--

## --seed-contents--

```html
<html>
--fcc-editable-region--
  
--fcc-editable-region--
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
