---
id: 5ef9b03c81a63668521804d2
title: Step 29
challengeType: 0
dashedName: step-29
---

# --description--

Msimbo wa orodha iliyopangwa (`ol`) ni sawa na orodha isiyopangwa, lakini vitu vya orodha katika orodha iliyopangwa hupewa nambari zinapoonyeshwa.

Chini ya kipengele cha `h3`, ongeza orodha iliyopangwa yenye vitu hivi vitatu vya orodha:

`flea treatment`
`thunder`
`other cats`

# --hints--

Kipengele chako cha `ol` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia hii: `<elementName>`.

```js
assert.exists(document.querySelector('ol'));
```

Kipengele chako cha `ol` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert.match(code, /<\/ol>/);
```

Kipengele cha `ol` kinapaswa kuwa juu ya lebo ya kufunga ya kipengele cha pili cha `section`. Umeviweka kwa mpangilio usio sahihi.

```js
assert.equal(document.querySelectorAll('main > section')[1]?.lastElementChild.nodeName, 'OL');
```

Vipengele vitatu vya `li` vinapaswa kupangwa ndani ya kipengele cha `ol`.

```js
assert.lengthOf(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ), 3
);
```

Unapaswa kuwa na vipengele vitatu vya `li` vyenye maandishi ya `flea treatment`, `thunder` na `other cats` kwa mpangilio wowote.

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText?.trim().replace(/\s+/g, ' ').toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

Unapaswa kuwa na kipengele kimoja tu cha `ol`.

```js 
assert.lengthOf([...document.querySelectorAll('ol')], 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```
