---
id: 5ef9b03c81a63668521804eb
title: Step 39
challengeType: 0
dashedName: step-39
---

# --description--

Kipengele cha `title` kinaamua kile vivinjari vinaonyesha kwenye upau wa kichwa au kichupo cha ukurasa.

Ongeza kipengele cha `title` ndani ya kipengele cha `head` ukitumia maandishi yaliyo hapa chini:

`CatPhotoApp`

# --hints--

Umefuta kipengele cha `head` au kinakosa lebo ya kufungua au lebo ya kufunga.

```js
assert.match(code, /\<head\>/);
assert.match(code, /\<\/head\>/);
```

Kipengele chako cha `title` kinapaswa kupangwa ndani ya kipengele cha `head`. Hakikisha umeongeza lebo ya kufungua na lebo ya kufunga kwa kipengele cha `title`.

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<head\>\<title\>.*\<\/title\>\<\/head\>/);
```

Kipengele chako cha `title` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert.match(code, /\<\/title\>/);
```

Maandishi ya kipengele chako cha `title` yanapaswa kuwa `CatPhotoApp`. Umeacha maandishi au kuna makosa ya tahajia.

```js
assert.equal(document.title?.toLowerCase(), 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <head>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
