---
id: 5ef9b03c81a63668521804e7
title: Step 35
challengeType: 0
dashedName: step-35
---

# --description--

Kipengele cha `footer` kinatumika kufafanua kijachini kwa hati au sehemu. Kijachini kawaida huwa na taarifa kuhusu mwandishi wa hati, data za hakimiliki, viungo vya masharti ya matumizi, taarifa za mawasiliano, na zaidi.

Baada ya kipengele cha `main`, ongeza kipengele cha `footer`.

# --hints--

Umefuta kipengele cha `main` au kinakosa lebo ya kufungua au lebo ya kufunga.

```js
assert.exists(document.querySelector('main'));
assert.match(code, /<\/main>/);
```

Kipengele chako cha `footer` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia ifuatayo: `<elementName>`.

```js
assert.exists(document.querySelector('footer'));
```

Kipengele chako cha `footer` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert.match(code, /<\/footer\>/);
```

Kipengele chako cha `footer` kinapaswa kuwa chini ya lebo ya kufunga ya kipengele cha `main`. Umeiweka mahali pengine.

```js
assert.equal(document.querySelector('main')?.nextElementSibling.nodeName, 'FOOTER');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
