---
id: 5ef9b03c81a63668521804e9
title: Step 37
challengeType: 0
dashedName: step-37
---

# --description--

Badilisha maandishi ya `freeCodeCamp.org` yaliyopo kuwa kiungo kwa kuyazunguka na kipengele cha nanga (`a`). Sifa ya `href` inapaswa kuwekwa kuwa `https://www.freecodecamp.org`.

# --hints--

Kipengele chako cha nanga (`a`) kinapaswa kupangwa ndani ya kipengele cha `p` ndani ya `footer`. Hakikisha unaongeza lebo ya kufungua na lebo ya kufunga kwa kipengele cha nanga (`a`).

```js
assert.isNotEmpty(document.querySelectorAll('footer > p > a'));
```

Kipengele chako cha nanga (`a`) kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert.exists(aElemClosingTags);
assert.lengthOf(aElemClosingTags, 4);
```

Kipengele chako cha nanga (`a`) kinapaswa kuwa na sifa ya `href` yenye thamani `https://www.freecodecamp.org`. Huenda umesahau sifa/thamani, au kuna makosa ya tahajia.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(nestedAnchor?.getAttribute('href').toLowerCase(), 'https://www.freecodecamp.org');
```

Maandishi ya kiungo yanapaswa kuwa `freeCodeCamp.org`. Huenda umesahau maandishi au kuna makosa ya tahajia.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(
  nestedAnchor?.innerText?.trim().toLowerCase(),
    'freecodecamp.org'
);
```

Baada ya kupanga ndani kipengele cha nanga (`a`), maudhui pekee ya kipengele cha `p` yanayoonekana kwenye kivinjari yanapaswa kuwa `No Copyright - freeCodeCamp.org`. Hakikisha upitie tena maandishi, nafasi, au alama za uandishi za `p` na kipengele cha nanga kilichopangwa ndani.

```js
const pText = document.querySelectorAll('footer > p')[0]?.innerText?.trim().toLowerCase().replace(/\s+/g, ' ');
assert.match(pText, /^no copyright - freecodecamp\.org$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```
