---
id: 5f07c98cdb9413cbd4b16750
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--

Ni wakati wa kuongeza sehemu mpya. Ongeza kipengele cha `section` cha pili chini ya kipengele cha `section` kilichopo.

# --hints--

Kipengele chako cha `section` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia hii: `<elementName>`.

```js
assert.isAtLeast(document.querySelectorAll('section').length, 2);
```

Unapaswa kuongeza lebo moja tu ya kufungua `section`. Tafadhali ondoa ziada zozote.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
```

Kipengele chako cha `section` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert.isAtLeast(code.match(/<\/section>/g)?.length, 2);
```

Unapaswa kuongeza lebo moja tu ya kufunga `section`. Tafadhali ondoa ziada zozote.

```js
assert.lengthOf(code.match(/<\/section>/g), 2);
```

Kipengele cha pili cha `section` hakipaswi kupangwa ndani ya kipengele cha kwanza cha `section`.

```js
const el = document.querySelector('main > section');
assert.isNotNull(el);
const childrenOf1stSection = [
  ...el.children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.isEmpty(foundElems);
```

Vipengele vyote viwili vya `section` vinapaswa kuwa kati ya lebo ya kufungua na lebo ya kufunga ya kipengele cha `main`.

```js
const el = document.querySelector('main');
assert.isNotNull(el);
const childrenOfMain = [...el.children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.lengthOf(foundElems, 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
