---
id: 68e4276483b0c71658aab1d0
title: Hatua 10
challengeType: 1
dashedName: step-10
---

# --description--

Kutakuwa na nyakati ambapo utahitaji kuunda mfuatano wa herufi mpya wenye herufi zilizo katika aina tofauti ya herufi.

Ili kuunda mfuatano wa herufi mpya wenye herufi zote zilizogeuzwa kuwa herufi kubwa, unaweza kutumia njia ya `toUpperCase()`.

Hapa kuna mfano:

```js
const sentence = "I love to code!";
sentence.toUpperCase(); // "I LOVE TO CODE!"
```

Unda kigezo kinachoitwa `upperCaseInput` na weka thamani yake kuwa matokeo ya kutumia njia ya `toUpperCase()` kwenye kigezo cha `cleanedInput`.

# --hints--

Unapaswa kuwa na kigezo cha `upperCaseInput`.

```js
assert.exists(upperCaseInput);
```

Kigezo chako cha `upperCaseInput` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(upperCaseInput);
```

Unapaswa kuweka thamani ya `cleanedInput.toUpperCase()` kwenye kigezo chako cha `upperCaseInput`.

```js
assert.equal(upperCaseInput, cleanedInput.toUpperCase());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

--fcc-editable-region--

--fcc-editable-region--
```
