---
id: 68e4276483b0c71658aab1d1
title: Hatua 12
challengeType: 1
dashedName: step-12
---

# --description--

Ikiwa unahitaji kuunda mfuatano mpya wa herufi ambapo herufi zote ziko katika herufi ndogo, unaweza kutumia njia ya `toLowerCase()`.  

Anza kwa kuunda kigezo kinachoitwa `lowerCaseInput` na ukaweke thamani ya matokeo ya kutumia njia ya `toLowerCase()` kwenye kigezo cha `cleanedInput`. 

Kisha andika mfuatano wa herufi `"Result of using the toLowerCase() method:"` kwenye konsoli.

Mwishowe, ongeza wito mwingine wa `console.log()` kwa kigezo cha `lowerCaseInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kuwa na kigezo cha `lowerCaseInput`.

```js
assert.exists(lowerCaseInput);
```

Kigezo chako cha `lowerCaseInput` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(lowerCaseInput);
```

Unapaswa kuweka thamani ya `cleanedInput.toLowerCase()` kwenye kigezo chako cha `lowerCaseInput`.

```js
assert.equal(lowerCaseInput, cleanedInput.toLowerCase());
```

Unapaswa kuandika mfuatano wa herufi `"Result of using the toLowerCase() method:"` kwenye konsoli.

```js
assert.equal(getLogs()[10], "Result of using the toLowerCase() method:");
```

Unapaswa kuandika kigezo cha `lowerCaseInput` kwenye konsoli.

```js
assert.equal(getLogs()[11], lowerCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);


const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
