---
id: 68e4276483b0c71658aab1cc
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Thamani yako ya `userInput` ina nafasi za ziada mwanzoni na mwishoni mwa mfuatano wa herufi. Ili kuunda mfuatano mpya wa herufi usio na nafasi hizo za ziada, unaweza kutumia njia ya `trim()`.  

Hapa kuna mfano:

```js
const sentence = "  Learning about the trim method.  ";
sentence.trim(); // "Learning about the trim method."
```

Unda kigezo kinachoitwa `cleanedInput` na ukaweke thamani ya matokeo ya kutumia njia ya `trim()` kwenye kigezo cha `userInput`.

# --hints--

Unapaswa kuwa na kigezo cha `cleanedInput`.

```js
assert.exists(cleanedInput);
```

Kigezo chako cha `cleanedInput` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(cleanedInput);
```

Unapaswa kuweka thamani ya `userInput.trim()` kwenye kigezo chako cha `cleanedInput`.

```js
assert.equal(cleanedInput, userInput.trim());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);
--fcc-editable-region--

--fcc-editable-region--
```
