---
id: 68e4276483b0c71658aab1ca
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utafanya mazoezi ya kutumia njia mbalimbali za mfuatano wa herufi zinazotumika kwa ajili ya kupanga mfuatano wa herufi.

Anza kwa kuunda kigezo kinachoitwa `userInput` na ukaweke thamani ifuatayo: `"   Hello World!   "`.

# --hints--

Unapaswa kuwa na kigezo cha `userInput`.

```js
assert.exists(userInput);
```

Kigezo chako cha `userInput` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(userInput);
```

Unapaswa kuweka thamani ya `"   Hello World!   "` katika kigezo chako cha `userInput`. Hakikisha kuna nafasi tatu kabla ya `H` na nafasi tatu baada ya `!`.

```js
assert.equal(userInput, "   Hello World!   ");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
