---
id: 68e4276483b0c71658aab1d3
title: Hatua 14
challengeType: 1
dashedName: step-14
---

# --description--

Kama unavyokumbuka kutoka mafundoni na warsha zilizopita, unaweza kutoa sehemu za mfuatano wa herufi kwa kutumia njia ya `slice()`.

Hapa kuna ukumbusho:

```js
const str = "abcd";
str.slice(0, 3); // "abc"
```

Unapaswa kubadilisha mgawo wa thamani wa `camelCasedVersion` uliopo kwa kubadilisha mfuatano wa herufi tupu na `lowercaseWord.slice(0, 5)`.

Sasa unapaswa kuona neno `"camel"` kwenye konsoli.

# --hints--

Kigezo chako cha `camelCasedVersion` hakipaswi kuwa mfuatano wa herufi tupu.

```js
assert.isNotEmpty(camelCasedVersion);
```

Unapaswa kuweka thamani ya `lowercaseWord.slice(0, 5)` kwa kigezo cha `camelCasedVersion`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5));
```

HUPASWI kuweka mfuatano wa herufi halisi `"camel"` kwa `camelCasedVersion`. Tafadhali soma tena maelekezo kwa njia sahihi.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camel(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = ""
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
