---
id: 68e4276483b0c71658aab1d2
title: Hatua 13
challengeType: 1
dashedName: step-13
---

# --description--

Kutakuwa na nyakati ambapo utahitaji kutumia njia nyingi kuunda mfuatano mpya wa herufi. Katika hatua chache zijazo, utatumia baadhi ya njia ulizojifunza hadi sasa kuunda mfuatano mpya wa herufi unaotumia mkataba wa camelcase.

Anza kwa kuunda kigezo kinachoitwa `lowercaseWord` na ukaweke thamani ya mfuatano wa herufi `"camelcase"`.

Chini ya kigezo hicho cha `lowercaseWord`, unda kigezo kingine kinachoitwa `camelCasedVersion` na ukaweke thamani ya mfuatano wa herufi tupu kwa sasa.

Chini ya kigezo cha `camelCasedVersion`, ongeza `console.log()` yenye mfuatano wa herufi `"Camel cased version:"` ikifuatiwa na `console.log()` kwa kigezo cha `camelCasedVersion`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `lowercaseWord`.

```js
assert.exists(lowercaseWord);
```

`lowercaseWord` yako inapaswa kuwa mfuatano wa herufi.

```js
assert.isString(lowercaseWord);
```

Unapaswa kuweka thamani ya mfuatano wa herufi `"camelcase"` kwa kigezo chako cha `lowercaseWord`.

```js
assert.equal(lowercaseWord, "camelcase");
```

Unapaswa kuwa na kigezo cha `camelCasedVersion`.

```js
assert.exists(camelCasedVersion);
```

Kigezo chako cha `camelCasedVersion` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(camelCasedVersion);
```

Unapaswa kuweka thamani ya mfuatano wa herufi tupu kwa kigezo chako cha `camelCasedVersion`.

```js
assert.isEmpty(camelCasedVersion);
```

Unapaswa kuandika mfuatano wa herufi `"Camel cased version:"` kwenye konsoli.

```js
assert.equal(getLogs()[12], "Camel cased version:");
```

Unapaswa kuandika kigezo cha `camelCasedVersion` kwenye konsoli.

```js
assert.equal(getLogs()[13], camelCasedVersion);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
