---
id: 68e4276483b0c71658aab1ce
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Kutakuwa na nyakati ambapo utataka tu kukata nafasi tupu mwanzoni mwa mfuatano wa herufi. Hapo ndipo njia ya `trimStart()` inavyokuwa na msaada.

Hapa kuna mfano:

```js
const sentence = "  I love learning.  ";
sentence.trimStart(); // "I love learning.  "
```

Kumbuka kuwa mfuatano wa herufi hauwezi kubadilika. Njia hii itaunda mfuatano mpya wa herufi wenye matokeo. Haita badilisha ule wa awali.

Tengeneza kigezo `trimmedStart` na ukaweke thamani matokeo ya kutumia njia ya `trimStart()` kwenye kigezo `userInput`.

# --hints--

Unapaswa kuwa na kigezo `trimmedStart`.

```js
assert.exists(trimmedStart);
```

Kigezo chako cha `trimmedStart` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(trimmedStart);
```

Unapaswa kuweka thamani ya `userInput.trimStart()` kwenye kigezo chako cha `trimmedStart`.

```js
assert.equal(trimmedStart, userInput.trimStart());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);
--fcc-editable-region--

--fcc-editable-region--
```
