---
id: 68e4276483b0c71658aab1cd
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Sasa ni wakati wa kuona matokeo.

Anza kwa kuongeza `console.log()` na mfuatano wa herufi `"Result of trimming whitespace from both ends:"`.

Kisha chini ya `console.log()`, ongeza `console.log()` nyingine yenye variable `cleanedInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kuandika mfuatano wa herufi `"Result of trimming whitespace from both ends:"` kwenye konsoli.

```js
assert.equal(getLogs()[2], "Result of trimming whitespace from both ends:");
```

Unapaswa kuandika variable `cleanedInput` kwenye konsoli.

```js
assert.equal(getLogs()[3], cleanedInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
--fcc-editable-region--

--fcc-editable-region--
```
