---
id: 68ed6079fe0e3e98b94a1b75
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Ifuatayo, ongeza `console.log()` na mfuatano wa herufi `"Original input:"`. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kuandika mfuatano wa herufi `"Original input:"` kwenye konsoli.

```js
assert.equal(getLogs()[0], "Original input:");
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
--fcc-editable-region--

--fcc-editable-region--
```
