---
id: 68f1179eaea3bb59ba434386
title: Hatua 8
challengeType: 1
dashedName: step-8
---

# --description--

Ili kukata nafasi tupu mwishoni mwa mfuatano wa herufi, unaweza kutumia njia ya `trimEnd()`.

Hapa kuna mfano:

```js
const sentence = "  I love learning.  ";
sentence.trimEnd(); // "  I love learning."
```

Tengeneza kigezo kinachoitwa `trimmedEnd` na weka thamani yake kuwa matokeo ya kutumia njia ya `trimEnd()` kwenye kigezo cha `userInput`. 

# --hints--

Unapaswa kuwa na kigezo cha `trimmedEnd`.

```js
assert.exists(trimmedEnd);
```

Kigezo chako cha `trimmedEnd` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(trimmedEnd);
```

Unapaswa kuweka thamani ya `userInput.trimEnd()` kwenye kigezo chako cha `trimmedEnd`.

```js
assert.equal(trimmedEnd, userInput.trimEnd());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

--fcc-editable-region--

--fcc-editable-region--
```
