---
id: 68f3f3d65d51d80ac0b412fc
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Sasa ni wakati wa kuona matokeo ya kukata nafasi zisizohitajika mwanzoni mwa mfuatano wa herufi.

Chini ya variable yako ya `trimmedStart`, ongeza `console.log()` yenye mfuatano wa herufi `"After using the trimStart() method, leading spaces removed:"`.

Kisha chini ya `console.log()`, ongeza `console.log()` nyingine yenye variable ya `trimmedStart`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kuandika mfuatano wa herufi `"After using the trimStart() method, leading spaces removed:"` kwenye konsoli.

```js
assert.equal(getLogs()[4], "After using the trimStart() method, leading spaces removed:");
```

Unapaswa kuandika variable ya `trimmedStart` kwenye konsoli.

```js
assert.equal(getLogs()[5], trimmedStart);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
--fcc-editable-region--

--fcc-editable-region--
```
