---
id: 68f3f890010c870e17bbe754
title: Hatua 11
challengeType: 1
dashedName: step-11
---

# --description--

Sasa ni wakati wa kuona matokeo.

Chini ya variable yako ya `upperCaseInput`, ongeza `console.log()` yenye mfuatano wa herufi `"Result of using the toUpperCase() method:"`. 

Kisha chini ya `console.log()` hiyo, ongeza `console.log()` nyingine yenye variable ya `upperCaseInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kuandika mfuatano wa herufi `"Result of using the toUpperCase() method:"` kwenye konsoli.

```js
assert.equal(getLogs()[8], "Result of using the toUpperCase() method:");
```

Unapaswa kuandika variable ya `upperCaseInput` kwenye konsoli.

```js
assert.equal(getLogs()[9], upperCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
--fcc-editable-region--

--fcc-editable-region--
```
