---
id: 68f3eeaa65614107ecd5f1f2
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Chini ya `console.log()` yako ya kwanza, ongeza `console.log()` nyingine ili kurekodi thamani ya mabadiliko ya `userInput` kwenye konsoli.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kurekodi mabadiliko ya `userInput` kwenye konsoli.

```js
assert.equal(getLogs()[1], userInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
--fcc-editable-region--

--fcc-editable-region--
```
