---
id: 68f3f5ca1d2a490c97f14f03
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Sasa ni wakati wa kuona matokeo ya kukata nafasi zisizohitajika mwishoni mwa mfuatano wa herufi.

Chini ya variable yako ya `trimmedEnd`, ongeza `console.log()` yenye mfuatano wa herufi `"After using the trimEnd() method, trailing spaces removed:"`. 

Kisha chini ya `console.log()`, ongeza tena `console.log()` yenye variable ya `trimmedEnd`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kuandika mfuatano wa herufi `"After using the trimEnd() method, trailing spaces removed:"` kwenye konsoli.

```js
assert.equal(getLogs()[6], "After using the trimEnd() method, trailing spaces removed:");
```

Unapaswa kuandika variable ya `trimmedEnd` kwenye konsoli.

```js
assert.equal(getLogs()[7], trimmedEnd);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
--fcc-editable-region--

--fcc-editable-region--
```
