---
id: 68f40925d810ae14b2d5a169
title: Hatua 16
challengeType: 1
dashedName: step-16
---

# --description--

Kwa sasa, variable yako ya `camelCasedVersion` ina thamani ya `camelC`. Lakini matokeo unayotaka ni `camelCase`. 

Ili kupata herufi zilizobaki za `ase`, unaweza kutumia njia ya `slice()` kwa thamani hasi ya kielezo. Kumbuka, kutumia thamani hasi za kielezo kunamwambia njia ya `slice()` kutoa kutoka mwisho. 

Hapa kuna mfano wa ukumbusho:

```js
const str = "abcd";
str.slice(-2); // cd
```

Sahihisha mgawo wa `camelCasedVersion` pia kuunganisha `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` na matokeo ya kutumia njia ya `slice()` kwenye `lowercaseWord`. Pitisha nambari `-3` kwa njia ya `slice()`.

Na kwa mabadiliko hayo ya mwisho, warsha yako itakuwa imekamilika!

# --hints--

Unapaswa kuunganisha `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` na `lowercaseWord.slice(-3)`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase() + lowercaseWord.slice(-3));
```

HUPASWI kuweka mfuatano wa herufi wa `"camelCase"` moja kwa moja kwenye `camelCasedVersion`. Tafadhali soma tena maelekezo kwa njia sahihi.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelCase(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

# --solutions--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";

const camelCasedVersion =
  lowercaseWord.slice(0, 5) +
  lowercaseWord[5].toUpperCase() +
  lowercaseWord.slice(-3);

console.log("Camel cased version:");
console.log(camelCasedVersion);
```
