---
id: 68f405c8c9df44132c8fe1bb
title: Hatua 15
challengeType: 1
dashedName: step-15
---

# --description--

Kama unavyokumbuka kutoka mafundoni ya awali, mkataba wa majina wa camelcase huanza na herufi ndogo kwa neno la kwanza, na kila neno linalofuata huanza na herufi kubwa ikifuatiwa na herufi ndogo.

Hapa kuna mifano ya kukukumbusha:

```md
freeCodeCamp
camelCase
learningIsFun
```

Neno la pili katika mabadiliko ya `lowercaseWord` ni `"case"`. Ili kupata upatikanaji wa `c` katika neno hilo, unaweza kutumia `lowercaseWord[5]`.

Sahihisha mgawo wa thamani wa `camelCasedVersion` ili utumie operator wa `+` kuunganisha `lowercaseWord.slice(0, 5)` na matokeo ya kutumia njia sahihi ya kubadilisha mfuatano wa herufi kuwa herufi kubwa kwenye `lowercaseWord[5]`.

Sasa unapaswa kuona `camelC` kwenye konsoli.

# --hints--

Unapaswa kuunganisha `lowercaseWord.slice(0, 5)` na `lowercaseWord[5].toUpperCase()`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase());
```

HUPASWI kuweka thamani ya mfuatano wa herufi `"camelC"` moja kwa moja kwa `camelCasedVersion`. Tafadhali soma tena maelekezo kwa njia sahihi.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelC(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5)
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
