---
id: 655cd899f8de09431eabb40c
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

归并排序是一种使用分治原理对数据集合进行排序的排序算法。也就是说，它将一个集合“分割”成更小的子部分，并通过独立排序这些子部分来“征服”它们，然后合并已排序的子部分。

Begin by creating a function called `merge_sort` that takes the parameter `array`. 此函数将处理排序数字列表的任务。 在函数体中使用 `pass` 关键字。

# --hints--

你应该声明一个名为 `merge_sort` 的函数，带有一个参数：`array`。不要忘记在函数体中使用 `pass` 关键字。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_args("array")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
