---
id: 6564aee9c077774ea49c3faf
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

归并排序算法执行三个操作：

- 将未排序的列表划分为子部分
- 在子部分中排序项目
- 合并已排序的子部分

上述过程递归进行，直到子部分合并为完整的排序序列。 让我们从划分序列开始。

首先，使用 `middle_point` 替换 `pass` 关键字。 使用整数除法运算符 (`//`) 将 `array` 列表的长度除以 2，并将结果分配给新的 `middle_point` 变量。 记得缩进你的代码。

# --hints--

函数主体中不应有 `pass`。

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).find_function("merge_sort").has_pass()`))
})
```

你应该声明一个 `middle_point` 变量并将 `len(array) // 2` 分配给它。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_stmt("middle_point = len(array) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
