---
id: 656639771fed09b5c0e8fe71
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

现在你已将 `array` 列表分成两个单独的列表，你将继续划分每个列表，直到每个元素都独立于其自己的列表中。 具有单个数字的列表始终是有序的。

为此，在你的函数内部递归地调用 `merge_sort`，并将 `left_part` 作为调用的参数传入。

# --hints--

你应该在函数体底部使用参数 `left_part` 调用 `merge_sort`。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("right_part = array[middle_point:]", "merge_sort(left_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
