---
id: 656680b0fc79f2c38a34d90e
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

Now it's time to sort and merge the lists (`left_part` and `right_part`) into the original list.

你可以通过比较两个列表中的元素并将较小的元素合并到主列表中来实现。 你将对 `left_part` 和 `right_part` 中的所有索引进行此比较。

创建三个变量：`left_array_index`、`right_array_index` 和 `sorted_index`，并将它们的值设置为 `0`。 这些变量将帮助您在排序过程中跟踪每个索引。

# --hints--

你应该在 `merge_sort` 函数中有一个名为 `left_array_index` 的变量

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_array_index")`)) 

})
```

The value of `left_array_index` should be `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("left_array_index").is_equivalent("left_array_index = 0")`))  

})
```

你应该在 `merge_sort` 函数中有一个名为 `right_array_index` 的变量

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_array_index")`))
})
```

The value of `right_array_index` should be `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_array_index").is_equivalent("right_array_index = 0")`))
})

```

你应该在 `merge_sort` 函数中有一个名为 `sorted_index` 的变量

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("sorted_index")`))
})
```

The value of `sorted_index` should be `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("sorted_index").is_equivalent("sorted_index = 0")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)
--fcc-editable-region--

--fcc-editable-region--
```
