---
id: 656627b47bd2d2a4afbc945d
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

在上一步中，你得到了中间点。 你可以使用它将 `array` 分成两部分，并将每个部分分配给新变量。 

使用切片语法提取`array`的左半部分并将其分配给名为`left_part`的变量。

# --hints--

`merge_sort` 函数中应该有一个名为 `left_part` 的变量。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_part")`)) 
})
```

你应该使用切片语法提取数组的左半部分并赋值给 `left_part` 变量。

```js
({
  test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_variable("left_part")
    cond.is_equivalent("left_part = array[:middle_point]") or cond.is_equivalent("left_part = array[0:middle_point]")
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
