---
id: 65670d1ef177e7e2b76d9528
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

你在前面的步骤中创建的 `if` 和 `else` 语句将元素分配给已排序的数组。

每个赋值给排序数组的元素都占据列表中的一个索引。因此，每次赋值后你必须移动到排序数组中的下一个索引。

在 `if`/`else` 块下面，但仍在 `while` 循环内，将 `sorted_index` 增加 `1`。

# --hints--

你应该在你的 `else` 子句后将 `sorted_index` 的当前值增加 `1`。

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_bodies()[0].find_body()[1]
    cond.is_equivalent("sorted_index += 1") or cond.is_equivalent("sorted_index = sorted_index + 1") or cond.is_equivalent("sorted_index = 1 + sorted_index")
    `))
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
--fcc-editable-region--
```
