---
id: 656707e11e671ed54c96f7ec
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

当 `if` 条件计算为 `True` 时，这意味着 `left_part` 列表中的元素小于 `right_part` 列表中与之比较的元素。

在这种情况下，你可以将位于 `left_array_index` 的元素放入排序数组中的 `left_part`。

在 `if` 块内，删除 `pass` 并将 `left_part[left_array_index]` 分配给 `array[sorted_index]`。

# --hints--

你应该在 `if` 语句主体中将 `pass` 替换为 `array[sorted_index] = left_part[left_array_index]`。

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_body().is_equivalent("array[sorted_index] = left_part[left_array_index]")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            pass

--fcc-editable-region--
```
