---
id: 656709e50ed928da35260276
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

在前面的步骤中，你将 `left_part` 中的元素分配给 `array` 列表，因为它较小。 但情况并非总是如此。 在一些比较情况下，右边的元素可能更小。 

创建一个 `else` 子句，当 `left_part` 中的元素不小于 `right_part` 中的元素时执行。

在 `else` 块内，将 `right_part[right_array_index]` 分配给 `array[sorted_index]`。

# --hints--

你应该在 `else` 子句的主体中包含 `array[sorted_index] = right_part[right_array_index]`。

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_bodies()[1].find_body()[0].is_equivalent("array[sorted_index] = right_part[right_array_index]")
`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
--fcc-editable-region--
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
--fcc-editable-region--
```
