---
id: 656702f8b4cbd8cbf0a433c6
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

接下来，你需要创建一个循环，只要 `left_part` 和 `right_part` 中都有剩余元素就继续。

为此，创建一个带有两个条件的 `while` 循环：一个条件查看 `left_array_index` 是否小于 `left_part` 的长度，另一个条件查看 `right_array_index` 是否小于 `right_part` 的长度。

# --hints--

你的 `while` 循环应查看 `left_array_index` 是否小于 `left_part` 的长度，以及 `right_array_index` 是否小于 `right_part` 的长度。

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_conditions()[0]

    cond.is_equivalent("left_array_index < len(left_part) and right_array_index < len(right_part)") or cond.is_equivalent("right_array_index < len(right_part) and left_array_index < len(left_part)")
    `))
}) 
```

在循环体中使用 `pass` 关键字。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[0].find_body().has_pass()`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
--fcc-editable-region--
    
--fcc-editable-region--
```
