---
id: 656710d1e0ec62253426db24
title: 步骤15
challengeType: 20
dashedName: step-15
---

# --description--

你创建的 `while` 循环将 `left_part` 中的一个元素与 `right_part` 中的另一个元素进行比较，然后将较小的元素添加到主要 `array` 列表中。

它将继续此操作，直到没有剩余元素可供比较。 但是 `left_part` 可能仍有元素，而 `right_part` 没有元素，反之亦然。

Create another `while` loop that runs when `left_array_index` is less than `len(left_part)`. In the next steps, you'll use it to copy the remaining elements in `left_part` into the `array` list.

# --hints--

你的 `while` 循环应具有以下条件：`left_array_index < len(left_part)`。 不要忘记在循环主体中使用 `pass` 关键字。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_conditions()[0].is_equivalent("left_array_index < len(left_part)")`))
})
```

你应该在 `while` 循环主体中添加 `pass` 关键字

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_bodies()[0].has_pass()`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
--fcc-editable-region--

--fcc-editable-region--
```
