---
id: 6569b20f829b7e69d43c232a
title: 步骤18
challengeType: 20
dashedName: step-18
---

# --description--

`while` 循环要做的最后一件事是移动到排序数组中的下一个索引。

Add `1` to the value of `sorted_index`.

# --hints--

You should increment the current value of `sorted_index` by one within the `while` loop. Note that the increment should be done after assigning the left part element to the sorted array

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[1].find_body()

    cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "sorted_index = 1 + sorted_index") 
    `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

--fcc-editable-region--
    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1

--fcc-editable-region--
```
