---
id: 6569b4e0bd29d17d4b53b16c
title: 步骤 19
challengeType: 20
dashedName: step-19
---

# --description--

现在，你将为 `right_part` 复制相同的 `while` 循环逻辑。

创建一个 `while` 循环，当 `right_array_index` 小于 `len(right_part)` 时运行。在 `while` 循环内，将 `right_part[right_array_index]` 赋值给 `array[sorted_index]`。最后，将 `right_array_index` 和 `sorted_index` 各自增加 1。

# --hints--

你的 `while` 循环应查看 `right_array_index` 是否小于 `len(right_part)`。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_conditions()[0].is_equivalent("right_array_index < len(right_part)")`)) 
})
```

你应该在循环体中将 `right_part[right_array_index]` 赋值给 `array[sorted_index]`。 

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_body().has_stmt("array[sorted_index] = right_part[right_array_index]")`))
})
```

你应该在 `while` 循环内将 `right_array_index` 增加 `1`。请注意，增加操作应在将右侧部分元素赋值到排序数组之后进行。

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = right_array_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = 1 + right_array_index")
    `))
}) 
```

你应该在 `while` 循环内将 `sorted_index` 的当前值增加 `1`。请注意，增加操作应在将右侧元素赋值到排序数组之后进行。

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = 1 + sorted_index")
    `))
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
