---
id: 6569b0a63423f65dd30888df
title: 步骤 16
challengeType: 20
dashedName: step-16
---

# --description--

移除 `pass` 关键字。 在 `while` 循环体中，将 `left_part[left_array_index]` 分配给 `array[sorted_index]`。

# --hints--

循环体中应该有 `array[sorted_index] = left_part[left_array_index]` 语句。 不要忘记移除循环体中的 `pass` 关键字。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_bodies()[0].find_body()[0].is_equivalent("array[sorted_index] = left_part[left_array_index]")`))
})
```

你不应该在 `while` 循环中使用 `pass`。

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_bodies()[0].has_pass()`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
        
--fcc-editable-region--
    while left_array_index < len(left_part):
        pass
--fcc-editable-region--
```
