---
id: 6577562501feabdf0984a184
title: 步骤11
challengeType: 20
dashedName: step-11
---

# --description--

在将 `left_part` 中的元素赋值给排序后的数组后，将 `left_array_index` 增加 `1`。

# --hints--

在将元素插入排序后的数组后，你应该将 `left_array_index` 的当前值加 `1`。

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_body()[1]

    cond.is_equivalent("left_array_index += 1") or cond.is_equivalent("left_array_index = left_array_index + 1") or cond.is_equivalent("left_array_index = 1 + left_array_index")
    `)) 
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
--fcc-editable-region--
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]

--fcc-editable-region--
```
