---
id: 6577254891048c8f2c19e961
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

在你的函数体底部，再次调用 `merge_sort()` 函数。这次，将 `right_part` 作为参数传递给函数调用。

# --hints--

你应该在你的函数体底部使用参数 `right_part` 调用 `merge_sort()`。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("merge_sort(left_part)", "merge_sort(right_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)

--fcc-editable-region--
```
