---
id: 657f3a661730891aa64f3568
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

使用切片语法提取 `array` 的右一半，并将其分配给变量 `right_part` 。

# --hints--

你应该有一个变量 `right_part`

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_part")`))
})
```

你应该使用切片语法提取数组的右半部分并赋值给 `right_part` 变量。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_part").is_equivalent("right_part = array[middle_point:]")`)) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
--fcc-editable-region--
```
