---
id: 657f59751d5a5c9b069ae0f3
title: 步骤21
challengeType: 20
dashedName: step-21
---

# --description--

正如你在之前的课程中学到的，你可以使用 `__name__` 变量来判断一个 Python 脚本是作为主程序运行还是作为模块（写在另一个 Python 文件中的代码）被导入。

如果 `__name__` 的值设置为 `'__main__'`，则表示当前脚本是主程序，而不是模块。

在这个项目中，你将使用当前脚本作为主程序。

创建一个 `if` 语句，查看 `__name__` 的值是否为 `'__main__'`。在该 `if` 语句内，创建一个名为 `numbers` 的列表，并将这些值赋值给它：`[4, 10, 6, 14, 2, 1, 8, 5]`。

# --hints--

`if` 语句中应该检查 `__name__` 的值是否为 `'__main__'`

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent("__name__ == '__main__'")
    `))
})
```

你应该有一个包含这些值的 `numbers` 列表：`[4, 10, 6, 14, 2, 1, 8, 5]`

```js
({
  test: () => assert(runPython(`_Node(_code).find_ifs()[0].find_body().find_variable("numbers").is_equivalent("numbers = [4, 10, 6, 14, 2, 1, 8, 5]")`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1

--fcc-editable-region--

--fcc-editable-region--
```
