---
id: 6a89c69cbe1bb69845992d8b
title: Construye un Generador de Rango de Números
challengeType: 27
dashedName: build-a-range-of-numbers-generator
---

# --description--

En este laboratorio, construirás una función `range_of_numbers` que usa recursión para generar una lista de números dentro de un rango especificado.

Por ejemplo:

- `range_of_numbers(3, 9)` debería devolver `[3, 4, 5, 6, 7, 8, 9]`.
- `range_of_numbers(5, 5)` debería devolver `[5]`.

Cumple las historias de usuario a continuación y haz que todas las pruebas pasen para completar el laboratorio.

**Historias de usuario:**

1. Debes crear una función llamada `range_of_numbers` con dos parámetros llamados `start_num` y `end_num`.
1. La función debe devolver una lista de enteros consecutivos que comienza con `start_num` y termina con `end_num`, incluyendo ambos valores.
1. El valor de `start_num` siempre será menor o igual que `end_num`.
1. Tu función debe usar recursión llamándose a sí misma.
1. Tu función no debe usar ciclos `for` o `while`, comprensiones ni la función `range()`.
1. Cuando `start_num` sea igual a `end_num`, la función debe devolver una lista que contenga `start_num`. Este es el caso base.
1. Para el caso recursivo, la función debe llamarse a sí misma con un argumento que se acerque al caso base, luego agregar el número actual a la lista devuelta.
1. La función no debe usar variables globales para almacenar la lista devuelta.

# --hints--

Debes tener una función llamada `range_of_numbers` con parámetros llamados `start_num` y `end_num`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert not _func.is_empty()
assert _func.has_args("start_num, end_num")
`) })
```

`range_of_numbers(1, 5)` debería devolver `[1, 2, 3, 4, 5]`.

```js
({ test: () => runPython(`assert range_of_numbers(1, 5) == [1, 2, 3, 4, 5]`) })
```

`range_of_numbers(6, 9)` debería devolver `[6, 7, 8, 9]`.

```js
({ test: () => runPython(`assert range_of_numbers(6, 9) == [6, 7, 8, 9]`) })
```

`range_of_numbers(4, 4)` debería devolver `[4]`.

```js
({ test: () => runPython(`assert range_of_numbers(4, 4) == [4]`) })
```

`range_of_numbers(10, 15)` debería devolver `[10, 11, 12, 13, 14, 15]`.

```js
({ test: () => runPython(`assert range_of_numbers(10, 15) == [10, 11, 12, 13, 14, 15]`) })
```

Tu función `range_of_numbers` debe llamarse a sí misma.

```js
({ test: () => runPython(`
assert _Node(_code).block_has_call("range_of_numbers", "range_of_numbers")
`) })
```

Tu función `range_of_numbers` no debe usar ciclos, comprensiones ni la función `range()`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert len(_func.find_for_loops()) == 0
assert len(_func.find_whiles()) == 0
assert len(_func.find_comps()) == 0
assert len(_func.find_calls("range")) == 0
`) })
```

Llamar a `range_of_numbers(1, 3)` no debe afectar una llamada posterior. Después de esa llamada, `range_of_numbers(6, 9)` aún debe devolver `[6, 7, 8, 9]`.

```js
({ test: () => runPython(`
range_of_numbers(1, 3)
assert range_of_numbers(6, 9) == [6, 7, 8, 9]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def range_of_numbers(start_num, end_num):
    if start_num == end_num:
        return [start_num]
    numbers = range_of_numbers(start_num, end_num - 1)
    numbers.append(end_num)
    return numbers
```
