---
id: 0095d0567b6817114989a1dd
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、JavaScriptで最短経路アルゴリズムを実装します。重み付きグラフのノード間の最短経路を計算し、距離と通った経路の両方を返す関数を書きます。

例えば、都市が異なる距離の道路でつながっているグラフがあるとき、ある都市から別の都市への最短ルートをアルゴリズムが見つけます。都市Aから都市Dへ移動したい場合、A ⇨ B ⇨ C ⇨ D（合計15km）と進むほうが、直接A ⇨ D（20km）に行くよりも短いかもしれません。

まず、`INF`という名前の変数を定義し、`Infinity`の値を代入してください。後で、2つのノード間に直接の接続がないことを示すために使います。

# --hints--

`INF`という名前の変数を定義してください。

```js
assert.isDefined(INF);
```

`INF`変数には`Infinity`の値を代入してください。

```js
assert.strictEqual(INF, Infinity);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
