---
id: 69c918841987bf16c1d1babc
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

グラフ内のノードの数をストアする必要があります。

`shortestPath` 関数の中で、`n` という名前の変数を作成し、`matrix.length` を代入してください。

# --hints--

`n` 関数の中に `shortestPath` という変数を作成してください。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+n\s*=/
);
```

`n` 変数には `matrix.length` の値を代入してください。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+n\s*=\s*matrix\.length/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {

}
--fcc-editable-region--
```
