---
id: 69c918841987bf16c1d1bac0
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

アルゴリズムが実行されるとき、すでに処理したノードを追跡して、再訪問しないようにする必要があります。

`visited` という名前の変数を作成し、`new Array(n).fill(false)` を代入してください。

# --hints--

`visited` という名前の変数を作成してください。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+visited\s*=/
);
```

`visited` 変数には `new Array(n).fill(false)` を代入してください。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*false\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
--fcc-editable-region--

--fcc-editable-region--
}
```
