---
id: 69c918841987bf16c1d1bac4
title: ステップ 13
challengeType: 1
dashedName: step-13
---

# --description--

条件が真の場合、現在のノードはこれまでに見つかった中で最良の未訪問候補です。追跡変数を更新してそれを反映させる必要があります。

`if` 文の中で、`minDistance` を `distances[nodeNo]` に、`current` を `nodeNo` に更新してください。

# --hints--

`minDistance` ブロック内で `distances[nodeNo]` を `if` に設定する必要があります。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /minDistance\s*=\s*distances\s*\[\s*nodeNo\s*\]/
);
```

`current` ブロック内で `nodeNo` を `if` に設定する必要があります。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /current\s*=\s*nodeNo/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        --fcc-editable-region--

        --fcc-editable-region--
      }
    }
  }
}
```
