---
id: 69c918841987bf16c1d1bac8
title: ステップ 17
challengeType: 1
dashedName: step-17
---

# --description--

距離を更新する前に、隣接ノードを検討に値するか確認する必要があります。実際にエッジが存在し（`distance !== INF`）、まだ訪問されていないことが条件です。

隣接ノードの`for`ループ内に、`if`が`distance`と厳密に等しくなく、かつ`INF`が未訪問であるかをチェックする`nodeNo`文を追加してください。

`if`ブロック内で、`newDistance`という変数を宣言し、`distances[current] + distance`を代入してください。これは現在のノードを経由して隣接ノードに到達するための合計コストです。

# --hints--

条件`if`を持つ`distance !== INF && !visited[nodeNo]`文を追加してください。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*distance\s*!==\s*INF\s*&&\s*!\s*visited\s*\[\s*\w+\s*\]\s*\)/
);
```

`newDistance`ブロック内で、値`distances[current] + distance`を持つ`if`を宣言してください。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+newDistance\s*=\s*distances\s*\[\s*current\s*\]\s*\+\s*distance/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
