---
id: 69c918841987bf16c1d1bac1
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

これからアルゴリズムを動かすメインループを追加します。このループはグラフ内の各ノードに対して1回ずつ実行され、そのたびに最も近い未訪問のノードを選びます。

`shortestPath` 関数の中に、`for` をループ変数として使い、`n` 回実行される `i` ループを追加してください。

ループ本体の最初で、`let` を使って、`minDistance` と `current` という2つの変数を宣言し、それぞれ `INF` と `-1` に設定してください。これらは各イテレーションで見つかった最も近い未訪問ノードを追跡します。

# --hints--

`for` の中に `shortestPath` ループがあり、`n` 回実行されているはずです。中で `i` に `let` を設定するのを忘れないでください。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+i\s*=\s*0\s*;\s*i\s*<\s*n\s*;\s*i\+\+\s*\)/
);
```

ループ内に `minDistance` 変数があり、`INF` に設定されているはずです。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+minDistance\s*=\s*INF/
);
```

ループ内に `current` 変数があり、`-1` に設定されているはずです。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+current\s*=\s*-1/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
--fcc-editable-region--

--fcc-editable-region--
}
```
