---
id: 69c937c73f23b1140bcd297a
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

`targetNode` パラメータは省略可能です。指定しない場合、関数は `startNode` からグラフ内の他のすべてのノードへの最短経路を計算します。

`targetNode` パラメータにデフォルト値として `null` を設定してください。

# --hints--

`targetNode` 関数の `shortestPath` パラメータは `null` をデフォルト値にしてください。

```js
assert.match(
  __helpers.removeJSComments(code),
  /targetNode\s*=\s*null/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode) {

}
--fcc-editable-region--
```
