---
id: 69c918841987bf16c1d1bac7
title: ステップ 16
challengeType: 1
dashedName: step-16
---

# --description--

現在のノードが訪問済みとしてマークされたので、すべての隣接ノードを見て、現在のノードを経由してより効率的に到達できるかどうかを確認する必要があります。

`visited[current] = true` の後に、`for` の隣接ノードループを追加して、`nodeNo` から `0` までの各 `n` を反復してください。これは最初のループと同じレベルの別の内側ループであることに注意してください。

ループの中で、`distance` という変数を宣言し、`matrix[current][nodeNo]` を代入してください。これにより、現在のノードとその隣接ノード間のエッジの重みが得られます。

# --hints--

現在のノードを訪問済みとしてマークした後に、隣接ノードの `for` ループを追加する必要があります。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(/
);
```

隣接ノードの `for` ループは、`0` から `n` まで反復する必要があります。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(\s*(var|let)\s+\w+\s*=\s*0\s*;\s*\w+\s*<\s*n\s*;\s*\w+\+\+\s*\)/
);
```

隣接ノードループの中で、`distance` 変数に `matrix[current][nodeNo]` を代入している必要があります。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distance\s*=\s*matrix\s*\[\s*current\s*\]\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
