---
id: 69c918841987bf16c1d1bac2
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

ループの中で、まだ訪問していない中で最も小さい既知の距離を持つノードを見つけるために、すべてのノードをスキャンする必要があります。

メインループの中に`for`ループを追加し、`nodeNo`をループ変数として`0`から`n`まで反復してください。

今はループの本体は空のままにしてください。

# --hints--

メインループの中に、`for`をループ変数とする2つ目（内側）の`nodeNo`ループがあるはずです。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+nodeNo\s*=\s*0\s*;\s*nodeNo\s*<\s*n\s*;\s*nodeNo\+\+\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

--fcc-editable-region--
  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;

  }
--fcc-editable-region--
}
```
