---
id: 69c918841987bf16c1d1babb
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

これから、すべての処理を行うメインの関数を作成します。

`shortestPath` という名前の関数を作成し、3つのパラメータ `matrix`、`startNode`、`targetNode` を受け取るようにしてください。

次のいくつかのステップで、この関数のロジックを構築していきます。

# --hints--

`shortestPath` という名前の関数を作成してください。

```js
assert.isFunction(shortestPath);
```

`shortestPath` 関数は `matrix`、`startNode`、`targetNode` のパラメータを持つ必要があります。

```js
const params = __helpers.getFunctionParams(shortestPath.toString());
const names = params.map(p => p.name);
assert.include(names, 'matrix');
assert.include(names, 'startNode');
assert.include(names, 'targetNode');
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--

--fcc-editable-region--
```
