---
id: 69c918841987bf16c1d1babe
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

開始ノードから自身への距離は常に `0` です。これを反映するように `distances` 配列を更新する必要があります。

`distances` 配列の後に、`distances[startNode]` を `0` に設定してください。

# --hints--

`distances[startNode]` を `0` に設定する必要があります。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /distances\s*\[\s*startNode\s*\]\s*=\s*0/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);

}
--fcc-editable-region--
```
