---
id: 69c918841987bf16c1d1babd
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

開始ノードから他のすべてのノードへの最短距離を追跡する必要があります。最初は、すべてのノードが無限に遠いと仮定します。

JavaScriptでは、`new Array(n).fill(value)`を使って値が事前に設定された配列を作成できます。例えば、`new Array(3).fill(0)`は`[0, 0, 0]`を作成します。

`shortestPath`関数の中で、`distances`という名前の変数を作成し、`new Array(n).fill(INF)`を代入してください。

# --hints--

`distances`の中に`shortestPath`という名前の変数を作成してください。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distances\s*=/
);
```

`distances`変数は`new Array(n).fill(INF)`に設定する必要があります。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*INF\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;

}
--fcc-editable-region--
```
