---
id: 69c918841987bf16c1d1bac3
title: ステップ 12
challengeType: 1
dashedName: step-12
---

# --description--

内側の`for`ループの中で、現在のノードが未訪問で、これまで見つけた中で最も近いかどうかを確認する必要があります。

`if`が未訪問であり、`nodeNo`が`distances[nodeNo]`より小さいかどうかをチェックする`minDistance`文を追加してください。

今は本体は空のままにしてください。

# --hints--

内側の`if`ループの中に、条件`!visited[nodeNo] && distances[nodeNo] < minDistance`を持つ`for`文があるはずです。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*!visited\s*\[\s*nodeNo\s*\]\s*&&\s*distances\s*\[\s*nodeNo\s*\]\s*<\s*minDistance\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
