---
id: 69c918841987bf16c1d1bac5
title: ステップ 14
challengeType: 1
dashedName: step-14
---

# --description--

すべてのノードをスキャンして最も近い未訪問のノードを見つけた後、有効なノードが見つからない場合の処理を行う必要があります。これは、残りのすべてのノードに到達できない場合に発生します。

内側の`for`ループの後、しかし外側のループの中にいる状態で、`if`が`current`と厳密に等しいかどうかをチェックする`-1`文を追加してください。それが真の場合、外側のループを早期に停止するために`break`ブロック内で`if`を使用してください。

# --hints--

`if`をチェックする`current === -1`文を追加してください。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)/
);
```

`break`ブロック内で`if`を使用してください。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)[\s\S]*?break/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
