---
id: 69c918841987bf16c1d1babf
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

各ノードに到達するまでの実際の経路も追跡する必要があります。経路のリストを保存し、各エントリーは通ったルートを表すノードのインデックスの配列になります。

最初は、各ノードの経路はそのノード自身だけを含みます。この構造は `Array.from()` を使って作成できます。

```js
Array.from({ length: n }, (_, i) => [i]);
```

これは `n` 個の要素を持つ配列を作成し、インデックス `i` の要素は `[i]` になります。

`shortestPath` 関数内で、`paths` という変数を作成し、`Array.from({ length: n }, (_, i) => [i])` を代入してください。

# --hints--

`paths` 関数内に `shortestPath` という変数を作成してください。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+paths\s*=/
);
```

`paths` 変数は `Array.from` を使って初期化してください。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*=\s*Array\.from/
);
```

`paths` 変数は `{ length: n }` の最初の引数として `Array.from` を使う必要があります。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\(\s*\{\s*length\s*:\s*n\s*\}/
);
```

`Array.from` の2番目の引数はコールバック `(_, i) => [i]` で、`_` は未使用の要素のプレースホルダー、`i` はインデックスです。戻り値は `[i]` でなければなりません。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\([\s\S]*?,\s*function\s*\(\s*\w+\s*,\s*(\w+)\s*\)\s*\{[\s\S]*?return\s*\[\s*\1\s*\]|Array\.from\s*\([\s\S]*?,\s*\(\s*\w+\s*,\s*\w+\s*\)\s*=>\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
