---
id: 69c918841987bf16c1d1bacb
title: ステップ 20
challengeType: 1
dashedName: step-20
---

# --description--

メインループが終了したら、どのノードの結果を表示するか決める必要があります。

特定の `targetNode` が指定されていれば、そのノードの結果だけを表示してください。そうでなければ、すべてのノードの結果を表示します。

JavaScriptでは、配列のすべてのインデックスをイテラブルとして取得するには `[...Array(n).keys()]` を使えます。

外側の `for` ループの後に、`targets` という名前の変数を作成してください。三項式を使い、`targetNode` が厳密に `null` でなければ `[targetNode]` を代入し、そうでなければ `[...Array(n).keys()]` を代入します。

# --hints--

メインループの後に `targets` という変数を作成してください。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+targets\s*=/
);
```

`targets` 変数は `targetNode !== null` をチェックする三項式を使うべきです。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /targets\s*=\s*targetNode\s*!==\s*null\s*\?/
);
```

ターゲットが指定されている場合、`targets` は `[targetNode]` でなければならず、そうでなければ `[...Array(n).keys()]` です。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /\?\s*\[\s*targetNode\s*\]\s*:\s*\[\s*\.\.\.\s*Array\s*\(\s*n\s*\)\s*\.keys\s*\(\s*\)\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
