---
id: 63507bcbfe4ede356e624395
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

现在你有了输入的值，你需要将其拆分成一个数字的数组。 使用 `.split()` 方法来实现这一点。

`.split()` 方法接收一个字符串，并将其拆分成一个字符串的数组。 你可以给它传入一个字母的字符串或一个正则表达式作为分隔符。 例如，`string.split(",")` 会在每个逗号处分割字符串并返回一个字符串的数组。

使用 `/,\s*/g` 正则通过逗号分割 `value` 字符串。 你可以根据值之间的空格数进行调整。 将该数组存储在一个 `array` 变量中。

# --hints--

你的 `calculate` 函数应该有一个 `array` 变量。

```js
assert.match(code.toString(), /array/);
```

你的 `calculate` 函数应该调用 `.split()` 方法，处理 `value` 变量。

```js
assert.match(code.toString(), /value\.split()/);
```

你的 `calculate` 函数应该将 `.split()` 方法的结果赋值给 `array` 变量。

```js
assert.match(code.toString(),  /array\s*=\s*value\.split()/);
```

你应该使用 `/,\s*/g` 作为 `split()` 方法的分隔符。

```js
assert.match(code.toString(), /value\.split\(\s*\/,\s*\\s*\*\s*\/g\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>

    <form onsubmit="calculate();">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
--fcc-editable-region--
const calculate = () => {
  const value = document.querySelector("#numbers").value;
}
--fcc-editable-region--
```
