---
id: 68d275dd800f404d22a07564
title: 깊이 우선 탐색 알고리즘 구현하기
challengeType: 27
dashedName: lab-depth-first-search
---

# --description--

이 실습에서는 <dfn>깊이 우선 탐색</dfn>이라는 그래프 순회 알고리즘을 구현할 것입니다.

<dfn>너비 우선 탐색</dfn>이 출발 노드에서 점차적으로 간선 길이를 늘려가며 검색하는 반면, <dfn>깊이 우선 탐색</dfn>은 먼저 한 경로를 가능한 한 깊게 따라 내려갑니다.

경로의 한 끝에 도달하면, 탐색은 방문하지 않은 간선 경로가 있는 마지막 노드로 되돌아가 계속 검색합니다.

너비 우선 탐색과 달리, 노드를 방문할 때마다 모든 이웃을 방문하지 않습니다. 대신, 먼저 이웃 중 하나를 방문하고 그 경로를 따라 더 이상 방문할 노드가 없을 때까지 계속 진행합니다.

이 알고리즘을 구현하려면 스택(마지막에 추가된 요소가 가장 먼저 제거되는 <dfn>후입선출</dfn> 원칙을 따르는 배열)을 사용하는 것이 좋습니다. 스택은 깊이 우선 탐색 알고리즘에서 유용한데, 이웃을 스택에 추가할 때 가장 최근에 추가된 이웃을 먼저 방문하고 스택에서 제거하기 때문입니다.

이 알고리즘의 간단한 출력은 주어진 노드에서 도달할 수 있는 노드들의 목록입니다. 따라서 방문한 노드를 추적하는 것도 필요합니다.

**목표**: 아래 사용자 스토리를 완성하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `dfs`라는 이름의 함수가 있어야 합니다.
1. `dfs` 함수는 두 개의 인수를 받아야 합니다:
   - 무방향 인접 행렬.
   - `0`와 `n - 1` 사이의 노드 숫자 레이블이며, `n`는 그래프의 총 노드 수입니다.
1. `dfs` 함수는 깊이 우선 탐색 알고리즘을 구현하여 전달된 노드에서 도달할 수 있는 모든 노드의 목록을 출력해야 합니다.

# --hints--

두 개의 인수를 받는 `dfs`라는 함수가 있어야 합니다.

```js
({ test: () => runPython(`
  import inspect
  assert inspect.isfunction(dfs)
  sig = inspect.signature(dfs)
  assert len(sig.parameters) == 2
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)`는 `1`, `2`, `3`, `0`이 포함된 목록을 반환해야 합니다.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)
  assert isinstance(result, list)
  assert len(result) == 4
  assert set(result) == {1, 2, 3, 0}
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)`는 `1`, `2`, `3`, `0`이 포함된 목록을 반환해야 합니다.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)
  assert isinstance(result, list)
  assert len(result) == 4
  assert set(result) == {3, 2, 1, 0}
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)`는 `[3]`을 반환해야 합니다.

```js
({ test: () => runPython(`
  assert dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3) == [3]
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)`는 `3`과 `2`이 포함된 목록을 반환해야 합니다.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)
  assert result == [3, 2] or result == [2, 3]
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)`는 `0`과 `1`이 포함된 목록을 반환해야 합니다.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)
  assert result == [0, 1] or result == [1, 0]
`) })
```

`dfs` 함수는 올바른 결과를 반환해야 합니다.

```js
({ test: () => runPython(`
  import random

  def solve(g, s):
      v, q = set(), [s]
      while q:
          u = q.pop()
          if u not in v:
              v.add(u)
              q.extend([i for i, x in enumerate(g[u]) if x == 1 and i not in v])
      return v

  random.seed(0)

  for _ in range(10):
      n = 6
      mat = [[0]*n for _ in range(n)]
      for i in range(n):
          for j in range(i+1, n):
              if random.random() > 0.5:
                  mat[i][j] = mat[j][i] = 1

      start = random.randint(0, n-1)
      assert set(dfs(mat, start)) == solve(mat, start), "Random graph check failed"
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def dfs(graph, root):
    stack = []
    temp_v = None
    visited = []
    temp_v_neighbors = []
    stack.append(root)
    while stack:
        temp_v = stack.pop()
        if temp_v not in visited:
            visited.append(temp_v)
            temp_v_neighbors = graph[temp_v]
            for n, is_neighbor in enumerate(temp_v_neighbors):
                if is_neighbor == 1:
                    stack.append(n)
    return visited
```
