---
id: 689f7996426c1534fa5bea66
title: Побудуйте перекладач результатів у гольфі
challengeType: 26
dashedName: build-a-golf-score-translator
---

# --description--

У грі в гольф кожна лунка має `par`, що означає середню кількість `strokes`, які гравець очікується зробити, щоб загнати м’яч у лунку і завершити гру. Залежно від того, наскільки ваші `par` відрізняються від `strokes`, існує різна прізвисько.

У цій лабораторній роботі ви напишете функцію, яка перетворює `par` і `strokes` у їхнє прізвисько.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні створити функцію з назвою `golfScore`.
1. `golfScore` має приймати два числові аргументи: пар для курсу та кількість зроблених штрихів.
1. `golfScore` має повертати рядок.
1. `golfScore` має повертати `"Hole-in-one!"`, якщо `strokes` дорівнює `1`.
1. `golfScore` має повертати `"Eagle"`, якщо `strokes` менше або дорівнює `par` мінус `2`.
1. `golfScore` має повертати `"Birdie"`, якщо `strokes` дорівнює `par` мінус `1`.
1. `golfScore` має повертати `"Par"`, якщо `strokes` дорівнює `par`.
1. `golfScore` має повертати `"Bogey"`, якщо `strokes` дорівнює `par` плюс `1`.
1. `golfScore` має повертати `"Double Bogey"`, якщо `strokes` дорівнює `par` плюс `2`.
1. `golfScore` має повертати `"Go Home!"`, якщо `strokes` більше або дорівнює `par` плюс `3`.

# --hints--

Ви повинні створити функцію з назвою `golfScore`.

```js
assert.isFunction(golfScore);
```

`golfScore` має приймати два параметри.

```js
assert.lengthOf(golfScore, 2);
```

`golfScore` має повертати рядок.

```js
assert.isString(golfScore(4,1)); 
```

`golfScore(1, 1)` має повертати рядок `Hole-in-one!`.

```js
assert.strictEqual(golfScore(1,1), 'Hole-in-one!');
```

`golfScore(3, 1)` має повертати рядок `Hole-in-one!`.

```js
assert.strictEqual(golfScore(3,1), 'Hole-in-one!');
```

`golfScore(4, 1)` має повертати рядок `Hole-in-one!`.

```js
assert.strictEqual(golfScore(4,1), 'Hole-in-one!');
```

`golfScore(5, 1)` має повертати рядок `Hole-in-one!`.

```js
assert.strictEqual(golfScore(5,1), 'Hole-in-one!');
```

`golfScore(4, 2)` має повертати рядок `Eagle`.

```js
assert.strictEqual(golfScore(4,2), 'Eagle');
```

`golfScore(5, 2)` має повертати рядок `Eagle`.

```js
assert.strictEqual(golfScore(5,2), 'Eagle');
```

`golfScore(3, 2)` має повертати рядок `Birdie`.

```js
assert.strictEqual(golfScore(3,2), 'Birdie');
```

`golfScore(4, 3)` має повертати рядок `Birdie`.

```js
assert.strictEqual(golfScore(4,3), 'Birdie');
```

`golfScore(5, 4)` має повертати рядок `Birdie`.

```js
assert.strictEqual(golfScore(5,4), 'Birdie');
```

`golfScore(3, 3)` має повертати рядок `Par`.

```js
assert.strictEqual(golfScore(3,3), 'Par');
```

`golfScore(4, 4)` має повертати рядок `Par`.

```js
assert.strictEqual(golfScore(4,4), 'Par');
```

`golfScore(5, 5)` має повертати рядок `Par`.

```js
assert.strictEqual(golfScore(5,5), 'Par');
```

`golfScore(3, 4)` має повертати рядок `Bogey`.

```js
assert.strictEqual(golfScore(3,4), 'Bogey');
```

`golfScore(4, 5)` має повертати рядок `Bogey`.

```js
assert.strictEqual(golfScore(4,5), 'Bogey');
```

`golfScore(5, 6)` має повертати рядок `Bogey`.

```js
assert.strictEqual(golfScore(5,6), 'Bogey');
```

`golfScore(3, 5)` має повертати рядок `Double Bogey`.

```js
assert.strictEqual(golfScore(3,5), 'Double Bogey');
```

`golfScore(4, 6)` має повертати рядок `Double Bogey`.

```js
assert.strictEqual(golfScore(4,6), 'Double Bogey');
```

`golfScore(5, 7)` має повертати рядок `Double Bogey`.

```js
assert.strictEqual(golfScore(5,7), 'Double Bogey');
```

`golfScore(3, 7)` має повертати рядок `Go Home!`.

```js
assert.strictEqual(golfScore(3,7), 'Go Home!');
```

`golfScore(4, 8)` має повертати рядок `Go Home!`.

```js
assert.strictEqual(golfScore(4,8), 'Go Home!');
```

`golfScore(5, 9)` має повертати рядок `Go Home!`.

```js
assert.strictEqual(golfScore(5,9), 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
