---
id: 56105e7b514f539506016a5e
title: العد تنازليًا باستخدام حلقة for
challengeType: 1
forumTopicId: 16808
dashedName: count-backwards-with-a-for-loop
---

# --description--

يمكن لحلقة for أن تعد تنازليًا أيضًا، طالما نحدد الشروط الصحيحة.

لكي ننقص بمقدار اثنين في كل تكرار، نحتاج إلى تغيير التهيئة، الشرط، والتعبير النهائي.

سنبدأ عند `i = 10` ونتكرر طالما `i > 0`. سننقص `i` بمقدار 2 في كل دورة باستخدام `i -= 2`.

```js
const ourArray = [];

for (let i = 10; i > 0; i -= 2) {
  ourArray.push(i);
}
```

سيحتوي `ourArray` الآن على `[10, 8, 6, 4, 2]`. دعنا نغير التهيئة والتعبير النهائي لنتمكن من العد تنازليًا بمقدار اثنين لإنشاء مصفوفة من الأعداد الفردية النازلة.

# --instructions--

ادفع الأعداد الفردية من 9 إلى 1 إلى `myArray` باستخدام حلقة `for`.

# --hints--

يجب أن تستخدم حلقة `for` لهذا.

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

يجب أن تستخدم طريقة المصفوفة `push`.

```js
assert(__helpers.removeJSComments(code).match(/myArray.push/));
```

يجب أن تساوي `myArray` القيمة `[9, 7, 5, 3, 1]`.

```js
assert.deepEqual(myArray, [9, 7, 5, 3, 1]);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 9; i > 0; i -= 2) {
  myArray.push(i);
}
```
