---
id: 56104e9e514f539506016a5c
title: تكرار الأعداد الفردية باستخدام حلقة For
challengeType: 1
forumTopicId: 18212
dashedName: iterate-odd-numbers-with-a-for-loop
---

# --description--

لا يجب أن تتكرر حلقات for واحدة تلو الأخرى. بتغيير `final-expression`، يمكننا العد بالأعداد الزوجية.

سنبدأ عند `i = 0` ونتكرر طالما `i < 10`. سنزيد `i` بمقدار 2 في كل تكرار باستخدام `i += 2`.

```js
const ourArray = [];

for (let i = 0; i < 10; i += 2) {
  ourArray.push(i);
}
```

ستحتوي `ourArray` الآن على `[0, 2, 4, 6, 8]`. لنغير `initialization` حتى نتمكن من العد بالأعداد الفردية.

# --instructions--

ادفع الأعداد الفردية من 1 إلى 9 إلى `myArray` باستخدام حلقة `for`.

# --hints--

يجب أن تستخدم حلقة `for` لهذا.

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

يجب أن تساوي `myArray` `[1, 3, 5, 7, 9]`.

```js
assert.deepEqual(myArray, [1, 3, 5, 7, 9]);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 1; i < 10; i += 2) {
  myArray.push(i);
}
```
