---
id: 56533eb9ac21ba0edf2244bd
title: تمرير القيم إلى الدوال باستخدام المعلمات
challengeType: 1
forumTopicId: 18254
dashedName: passing-values-to-functions-with-arguments
---

# --description--

<dfn>المعلمات</dfn> هي متغيرات تعمل كعناصر نائب للقيم التي ستُدخل إلى الدالة عند استدعائها. عند تعريف دالة، عادةً ما تُعرف مع واحدة أو أكثر من المعلمات. القيم الفعلية التي تُدخل (أو <dfn>"تُمرر"</dfn>) إلى الدالة عند استدعائها تُعرف باسم <dfn>المعلمات</dfn>.

إليك دالة تحتوي على معلمتين، `param1` و `param2`:

```js
function testFun(param1, param2) {
  console.log(param1, param2);
}
```

بعدها يمكننا استدعاء `testFun` هكذا: `testFun("Hello", "World");`. لقد مررنا معلمتين من نوع سلسلة نصية، `Hello` و `World`. داخل الدالة، ستساوي `param1` السلسلة النصية `Hello` و `param2` ستساوي السلسلة النصية `World`. لاحظ أنه يمكنك استدعاء `testFun` مرة أخرى بمعلمات مختلفة وستأخذ المعلمات قيم المعلمات الجديدة.

# --instructions--

<ol><li>أنشئ دالة باسم <code>functionWithArgs</code> تقبل معلمتين وتُخرج مجموعهما إلى وحدة التحكم للمطور.</li><li>استدعِ الدالة بمعلمتين رقميتين.</li></ol>

# --before-each--

```js
var logOutput;
```

# --hints--

يجب أن تكون `functionWithArgs` دالة.

```js
assert(typeof functionWithArgs === 'function');
```

يجب أن تُخرج `functionWithArgs(1,2)` القيمة `3`.

```js
console.log = function (message) {
  logOutput = message
};
functionWithArgs(1, 2);
assert(logOutput == 3);
```

يجب أن تُخرج `functionWithArgs(7,9)` القيمة `16`.

```js
console.log = function (message) {
  logOutput = message
};
functionWithArgs(7, 9);
assert(logOutput == 16);
```

يجب استدعاء `functionWithArgs` بمعلمتين رقميتين بعد تعريفها.

```js
assert(
  /functionWithArgs\([-+]?\d*\.?\d*,[-+]?\d*\.?\d*\)/.test(
    __helpers.removeJSComments(code).replace(/\s/g, '')
  )
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function functionWithArgs(a, b) {
  console.log(a + b);
}
functionWithArgs(10, 5);
```
