---
id: 56533eb9ac21ba0edf2244d5
title: المقارنة باستخدام معامل أكبر من أو يساوي
challengeType: 1
forumTopicId: 16785
dashedName: comparison-with-the-greater-than-or-equal-to-operator
---

# --description--

يقارن معامل أكبر من أو يساوي (`>=`) قيم رقمين. إذا كان الرقم على اليسار أكبر من أو يساوي الرقم على اليمين، يُرجع `true`. وإلا، يُرجع `false`.

مثل معامل المساواة، يقوم معامل أكبر من أو يساوي بتحويل أنواع البيانات أثناء المقارنة.

**أمثلة**

```js
6   >=  6  // true
7   >= '3' // true
2   >=  3  // false
'7' >=  9  // false
```

# --instructions--

أضف معامل أكبر من أو يساوي إلى الأسطر المشار إليها بحيث تكون جمل الإرجاع منطقية.

# --hints--

`testGreaterOrEqual(0)` يجب أن تُرجع السلسلة النصية `Less than 10`

```js
assert(testGreaterOrEqual(0) === 'Less than 10');
```

`testGreaterOrEqual(9)` يجب أن تُرجع السلسلة النصية `Less than 10`

```js
assert(testGreaterOrEqual(9) === 'Less than 10');
```

`testGreaterOrEqual(10)` يجب أن تُرجع السلسلة النصية `10 or Over`

```js
assert(testGreaterOrEqual(10) === '10 or Over');
```

`testGreaterOrEqual(11)` يجب أن تُرجع السلسلة النصية `10 or Over`

```js
assert(testGreaterOrEqual(11) === '10 or Over');
```

`testGreaterOrEqual(19)` يجب أن تُرجع السلسلة النصية `10 or Over`

```js
assert(testGreaterOrEqual(19) === '10 or Over');
```

`testGreaterOrEqual(100)` يجب أن تُرجع السلسلة النصية `20 or Over`

```js
assert(testGreaterOrEqual(100) === '20 or Over');
```

`testGreaterOrEqual(21)` يجب أن تُرجع السلسلة النصية `20 or Over`

```js
assert(testGreaterOrEqual(21) === '20 or Over');
```

يجب أن تستخدم معامل `>=` مرتين على الأقل

```js
assert(__helpers.removeJSComments(code).match(/val\s*>=\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testGreaterOrEqual(val) {
  if (val) {  // Change this line
    return "20 or Over";
  }

  if (val) {  // Change this line
    return "10 or Over";
  }

  return "Less than 10";
}

testGreaterOrEqual(10);
```

# --solutions--

```js
function testGreaterOrEqual(val) {
  if (val >= 20) {  // Change this line
    return "20 or Over";
  }

  if (val >= 10) {  // Change this line
    return "10 or Over";
  }

  return "Less than 10";
}
```
