---
id: 56533eb9ac21ba0edf2244ad
title: إنقاص رقم باستخدام JavaScript
challengeType: 1
forumTopicId: 17558
dashedName: decrement-a-number-with-javascript
---

# --description--

يمكنك بسهولة <dfn>إنقاص</dfn> أو تقليل قيمة متغير بمقدار واحد باستخدام العامل `--`.

```js
i--;
```

يعادل

```js
i = i - 1;
```

**ملاحظة:** تصبح السطر بأكمله `i--;`، مما يلغي الحاجة إلى علامة المساواة.

# --instructions--

غيّر الشفرة لاستخدام العامل `--` على `myVar`.

# --hints--

يجب أن تساوي `myVar` القيمة `10`.

```js
assert(myVar === 10);
```

يجب تغيير `myVar = myVar - 1;`.

```js
assert(!__helpers.removeJSComments(code).match(/myVar\s*=\s*myVar\s*[-]\s*1.*?;?/));
```

لا يجب تعيين `myVar` إلى `10`.

```js
assert(!__helpers.removeJSComments(code).match(/myVar\s*=\s*10.*?;?/));
```

يجب استخدام العامل `--` على `myVar`.

```js
assert(/[-]{2}\s*myVar|myVar\s*[-]{2}/.test(__helpers.removeJSComments(code)));
```

لا يجب تغيير الشفرة فوق التعليق المحدد.

```js
assert(/let myVar = 11;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let myVar = 11;

// Only change code below this line
myVar = myVar - 1;
```

# --solutions--

```js
let myVar = 11;
myVar--;
```
