---
id: 56533eb9ac21ba0edf2244c3
title: التعيين مع قيمة إرجاع
challengeType: 1
forumTopicId: 16658
dashedName: assignment-with-a-returned-value
---

# --description--

إذا تذكرت من مناقشتنا حول <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/storing-values-with-the-assignment-operator" target="_blank" rel="noopener noreferrer nofollow">تخزين القيم باستخدام معامل التعيين</a>، كل شيء على يمين علامة المساواة يُحل قبل تعيين القيمة. هذا يعني أننا نستطيع أخذ قيمة الإرجاع من دالة وتعيينها إلى متغير.

افترض أننا عرّفنا دالة `sum` التي تجمع رقمين معًا.

```js
ourSum = sum(5, 12);
```

استدعاء دالة `sum` بالمعلمتين `5` و `12` ينتج قيمة إرجاع `17`. تُعيَّن هذه القيمة إلى المتغير `ourSum`.

# --instructions--

استدعِ الدالة `processArg` بمعلمة `7` وعيّن قيمة الإرجاع إلى المتغير `processed`.

# --hints--

يجب أن تكون قيمة `processed` هي `2`

```js
assert(processed === 2);
```

يجب أن تعيّن `processArg` إلى `processed`

```js
assert(/processed\s*=\s*processArg\(\s*7\s*\)/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
let processed = 0;

function processArg(num) {
  return (num + 3) / 5;
}

// Only change code below this line

```

# --solutions--

```js
var processed = 0;

function processArg(num) {
  return (num + 3) / 5;
}

processed = processArg(7);
```
