---
id: 56533eb9ac21ba0edf2244c4
title: نمط الإرجاع المبكر للدوال
challengeType: 1
forumTopicId: 18272
dashedName: return-early-pattern-for-functions
---

# --description--

عند الوصول إلى جملة `return`، يتوقف تنفيذ الدالة الحالية ويعود التحكم إلى موقع الاستدعاء.

**مثال**

```js
function myFun() {
  console.log("Hello");
  return "World";
  console.log("byebye")
}
myFun();
```

سيعرض ما سبق السلسلة النصية `Hello` في وحدة التحكم، ويرجع السلسلة النصية `World`. لن تظهر السلسلة النصية `byebye` في وحدة التحكم أبدًا، لأن الدالة تخرج عند جملة `return`.

# --instructions--

عدّل الدالة `abTest` بحيث إذا كان `a` أو `b` أقل من `0` تخرج الدالة فورًا بقيمة `undefined`.

**تلميح**  
تذكر أن <a href="https://www.freecodecamp.org/learn/javascript-algorithms-and-data-structures/basic-javascript/understanding-uninitialized-variables" target="_blank" rel="noopener noreferrer nofollow"><code>undefined</code> هي كلمة مفتاحية</a> وليست سلسلة نصية.

# --hints--

`abTest(2, 2)` يجب أن تُرجع رقمًا

```js
assert(typeof abTest(2, 2) === 'number');
```

`abTest(2, 2)` يجب أن تُرجع `8`

```js
assert(abTest(2, 2) === 8);
```

`abTest(-2, 2)` يجب أن تُرجع `undefined`

```js
assert(abTest(-2, 2) === undefined);
```

`abTest(2, -2)` يجب أن تُرجع `undefined`

```js
assert(abTest(2, -2) === undefined);
```

`abTest(2, 8)` يجب أن تُرجع `18`

```js
assert(abTest(2, 8) === 18);
```

`abTest(3, 3)` يجب أن تُرجع `12`

```js
assert(abTest(3, 3) === 12);
```

`abTest(0, 0)` يجب أن تُرجع `0`

```js
assert(abTest(0, 0) === 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function abTest(a, b) {
  // Only change code below this line



  // Only change code above this line

  return Math.round(Math.pow(Math.sqrt(a) + Math.sqrt(b), 2));
}

abTest(2,2);
```

# --solutions--

```js
function abTest(a, b) {
  if(a < 0 || b < 0) {
    return undefined;
  }
  return Math.round(Math.pow(Math.sqrt(a) + Math.sqrt(b), 2));
}
```
