---
id: 56533eb9ac21ba0edf2244ba
title: فهم ثبات السلسلة النصية
challengeType: 1
forumTopicId: 18331
dashedName: understand-string-immutability
---

# --description--

في JavaScript، قيم `String` <dfn>ثابتة</dfn>، مما يعني أنه لا يمكن تعديلها بعد إنشائها.

على سبيل المثال، الشفرة التالية ستُنتج خطأ لأن الحرف `B` في السلسلة النصية `Bob` لا يمكن تغييره إلى الحرف `J`:

```js
let myStr = "Bob";
myStr[0] = "J";
```

لاحظ أن هذا لا يعني *عدم* إمكانية إعادة تعيين `myStr`. الطريقة الوحيدة لتغيير `myStr` هي تعيينه بقيمة جديدة، مثل هذا:

```js
let myStr = "Bob";
myStr = "Job";
```

# --instructions--

صحح التعيين إلى `myStr` ليحتوي على القيمة النصية `Hello World` باستخدام الطريقة الموضحة في المثال أعلاه.

# --hints--

يجب أن تكون قيمة `myStr` هي السلسلة النصية `Hello World`.

```js
assert(myStr === 'Hello World');
```

لا يجب عليك تغيير الشفرة فوق التعليق المحدد.

```js
assert(/myStr = "Jello World"/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
let myStr = "Jello World";

// Only change code below this line
myStr[0] = "H"; // Change this line
// Only change code above this line
```

# --solutions--

```js
let myStr = "Jello World";
myStr = "Hello World";
```
