---
id: 56533eb9ac21ba0edf2244da
title: تقديم جمل else
challengeType: 1
forumTopicId: 18207
dashedName: introducing-else-statements
---

# --description--

عندما يكون شرط جملة `if` صحيحًا، تُنفذ كتلة الشفرة التي تليها. ماذا عن عندما يكون هذا الشرط خاطئًا؟ عادةً لا يحدث شيء. باستخدام جملة `else`، يمكن تنفيذ كتلة شفرة بديلة.

```js
if (num > 10) {
  return "Bigger than 10";
} else {
  return "10 or Less";
}
```

# --instructions--

ادمج جمل `if` في جملة واحدة `if/else`.

# --hints--

يجب أن يكون لديك جملة `if` واحدة فقط في المحرر

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

يجب أن تستخدم جملة `else`

```js
assert(/else/g.test(__helpers.removeJSComments(code)));
```

يجب أن تُرجع `testElse(4)` السلسلة النصية `5 or Smaller`

```js
assert(testElse(4) === '5 or Smaller');
```

يجب أن تُرجع `testElse(5)` السلسلة النصية `5 or Smaller`

```js
assert(testElse(5) === '5 or Smaller');
```

يجب أن تُرجع `testElse(6)` السلسلة النصية `Bigger than 5`

```js
assert(testElse(6) === 'Bigger than 5');
```

يجب أن تُرجع `testElse(10)` السلسلة النصية `Bigger than 5`

```js
assert(testElse(10) === 'Bigger than 5');
```

لا يجب أن تغيّر الشفرة أعلاه أو أسفل التعليقات المحددة.

```js
assert(/let result = "";/.test(__helpers.removeJSComments(code)) && /return result;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function testElse(val) {
  let result = "";
  // Only change code below this line

  if (val > 5) {
    result = "Bigger than 5";
  }

  if (val <= 5) {
    result = "5 or Smaller";
  }

  // Only change code above this line
  return result;
}

testElse(4);
```

# --solutions--

```js
function testElse(val) {
  let result = "";
  if(val > 5) {
    result = "Bigger than 5";
  } else {
    result = "5 or Smaller";
  }
  return result;
}
```
