---
id: 56533eb9ac21ba0edf2244db
title: تقديم جمل Else If
challengeType: 1
forumTopicId: 18206
dashedName: introducing-else-if-statements
---

# --description--

إذا كان لديك شروط متعددة يجب معالجتها، يمكنك ربط جمل `if` مع جمل `else if`.

```js
if (num > 15) {
  return "Bigger than 15";
} else if (num < 5) {
  return "Smaller than 5";
} else {
  return "Between 5 and 15";
}
```

# --instructions--

حوّل المنطق لاستخدام جمل `else if`.

# --hints--

يجب أن يكون لديك على الأقل جملتان `else`

```js
assert(__helpers.removeJSComments(code).match(/else/g).length > 1);
```

يجب أن يكون لديك على الأقل جملتان `if`

```js
assert(__helpers.removeJSComments(code).match(/if/g).length > 1);
```

يجب أن تحتوي كل كتلة شفرة `if else` على أقواس متعرجة مفتوحة ومغلقة.

```js
assert(
  __helpers.removeJSComments(code).match(
    /if\s*\((.+)\)\s*\{[\s\S]+\}\s*else\s+if\s*\((.+)\)\s*\{[\s\S]+\}\s*else\s*\{[\s\S]+\s*\}/
  )
);
```

يجب أن تُرجع الدالة `testElseIf(0)` السلسلة النصية `Smaller than 5`

```js
assert(testElseIf(0) === 'Smaller than 5');
```

يجب أن تُرجع الدالة `testElseIf(5)` السلسلة النصية `Between 5 and 10`

```js
assert(testElseIf(5) === 'Between 5 and 10');
```

يجب أن تُرجع الدالة `testElseIf(7)` السلسلة النصية `Between 5 and 10`

```js
assert(testElseIf(7) === 'Between 5 and 10');
```

يجب أن تُرجع الدالة `testElseIf(10)` السلسلة النصية `Between 5 and 10`

```js
assert(testElseIf(10) === 'Between 5 and 10');
```

يجب أن تُرجع الدالة `testElseIf(12)` السلسلة النصية `Greater than 10`

```js
assert(testElseIf(12) === 'Greater than 10');
```

# --seed--

## --seed-contents--

```js
function testElseIf(val) {
  if (val > 10) {
    return "Greater than 10";
  }

  if (val < 5) {
    return "Smaller than 5";
  }

  return "Between 5 and 10";
}

testElseIf(7);
```

# --solutions--

```js
function testElseIf(val) {
  if(val > 10) {
    return "Greater than 10";
  } else if(val < 5) {
    return "Smaller than 5";
  } else {
    return "Between 5 and 10";
  }
}
```
