---
id: 56533eb9ac21ba0edf2244cd
title: الوصول إلى المصفوفات المسكنة
challengeType: 1
forumTopicId: 16160
dashedName: accessing-nested-arrays
---

# --description--

كما رأينا في الأمثلة السابقة، يمكن أن تحتوي الكائنات على كائنات مسكنة ومصفوفات مسكنة. وبالمثل مع الوصول إلى الكائنات المسكنة، يمكن ربط ترميز الأقواس للوصول إلى المصفوفات المسكنة.

إليك مثالًا على كيفية الوصول إلى مصفوفة مسكنة:

```js
const ourPets = [
  {
    animalType: "cat",
    names: [
      "Meowzer",
      "Fluffy",
      "Kit-Cat"
    ]
  },
  {
    animalType: "dog",
    names: [
      "Spot",
      "Bowser",
      "Frankie"
    ]
  }
];

ourPets[0].names[1];
ourPets[1].names[0];
```

`ourPets[0].names[1]` ستكون السلسلة النصية `Fluffy`، و`ourPets[1].names[0]` ستكون السلسلة النصية `Spot`.

# --instructions--

باستخدام ترميز النقطة وترميز الأقواس، عرّف المتغير `secondTree` ليكون العنصر الثاني في مصفوفة `list` من الكائن الثاني في مصفوفة `myPlants`.

# --hints--

يجب أن يساوي `secondTree` السلسلة النصية `pine`.

```js
assert(secondTree === 'pine');
```

يجب أن يستخدم شفرتك ترميز النقطة وترميز الأقواس للوصول إلى `myPlants`.

```js
assert(/=\s*myPlants\[1\].list\[1\]/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myPlants = [
  {
    type: "flowers",
    list: [
      "rose",
      "tulip",
      "dandelion"
    ]
  },
  {
    type: "trees",
    list: [
      "fir",
      "pine",
      "birch"
    ]
  }
];

const secondTree = "";
```

# --solutions--

```js
const myPlants = [
  {
    type: "flowers",
    list: [
      "rose",
      "tulip",
      "dandelion"
    ]
  },
  {
    type: "trees",
    list: [
      "fir",
      "pine",
      "birch"
    ]
  }
];

const secondTree = myPlants[1].list[1];
```
