---
id: 56533eb9ac21ba0edf2244c0
title: النطاق العام مقابل النطاق المحلي في الدوال
challengeType: 1
forumTopicId: 18194
dashedName: global-vs--local-scope-in-functions
---

# --description--

يمكن أن يكون لديك متغيرات <dfn>محلية</dfn> و<dfn>عامة</dfn> بنفس الاسم. عند حدوث ذلك، يتفوق المتغير المحلي على المتغير العام.

في هذا المثال:

```js
const someVar = "Hat";

function myFun() {
  const someVar = "Head";
  return someVar;
}
```

الدالة `myFun` سترجع السلسلة النصية `Head` لأن النسخة المحلية من المتغير موجودة.

# --instructions--

أضف متغيرًا محليًا إلى الدالة `myOutfit` لتجاوز قيمة `outerWear` بالسلسلة النصية `sweater`.

# --hints--

لا يجب أن تغير قيمة المتغير العام `outerWear`.

```js
assert(outerWear === 'T-Shirt');
```

يجب أن ترجع `myOutfit` السلسلة النصية `sweater`.

```js
assert(myOutfit() === 'sweater');
```

لا يجب أن تغير جملة الإرجاع.

```js
assert(/return outerWear/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const outerWear = "T-Shirt";

function myOutfit() {
  // Only change code below this line

  // Only change code above this line
  return outerWear;
}

myOutfit();
```

# --solutions--

```js
const outerWear = "T-Shirt";
function myOutfit() {
  const outerWear = "sweater";
  return outerWear;
}
```
