---
id: 56533eb9ac21ba0edf2244e1
title: التسكين في حلقات التكرار
challengeType: 1
forumTopicId: 18248
dashedName: nesting-for-loops
---

# --description--

إذا كان لديك مصفوفة متعددة الأبعاد، يمكنك استخدام نفس المنطق كما في نقطة التوقف السابقة للتكرار عبر المصفوفة وأي مصفوفات فرعية. إليك مثالًا:

```js
const arr = [
  [1, 2], [3, 4], [5, 6]
];

for (let i = 0; i < arr.length; i++) {
  for (let j = 0; j < arr[i].length; j++) {
    console.log(arr[i][j]);
  }
}
```

هذا يعرض كل عنصر فرعي في `arr` واحدًا تلو الآخر. لاحظ أننا في الحلقة الداخلية نتحقق من `.length` لـ `arr[i]`، لأن `arr[i]` هي مصفوفة بحد ذاتها.

# --instructions--

عدّل الدالة `multiplyAll` بحيث تُرجع حاصل الضرب لجميع الأرقام في المصفوفات الفرعية لـ `arr`.

# --hints--

`multiplyAll([[1], [2], [3]])` يجب أن تُرجع `6`

```js
assert(multiplyAll([[1], [2], [3]]) === 6);
```

`multiplyAll([[1, 2], [3, 4], [5, 6, 7]])` يجب أن تُرجع `5040`

```js
assert(
  multiplyAll([
    [1, 2],
    [3, 4],
    [5, 6, 7]
  ]) === 5040
);
```

`multiplyAll([[5, 1], [0.2, 4, 0.5], [3, 9]])` يجب أن تُرجع `54`

```js
assert(
  multiplyAll([
    [5, 1],
    [0.2, 4, 0.5],
    [3, 9]
  ]) === 54
);
```

# --seed--

## --seed-contents--

```js
function multiplyAll(arr) {
  let product = 1;
  // Only change code below this line

  // Only change code above this line
  return product;
}

multiplyAll([[1, 2], [3, 4], [5, 6, 7]]);
```

# --solutions--

```js
function multiplyAll(arr) {
  let product = 1;
  for (let i = 0; i < arr.length; i++) {
    for (let j = 0; j < arr[i].length; j++) {
      product *= arr[i][j];
    }
  }
  return product;
}
```
