---
id: 56533eb9ac21ba0edf2244d4
title: المقارنة باستخدام معامل أكبر من
challengeType: 1
forumTopicId: 16786
dashedName: comparison-with-the-greater-than-operator
---

# --description--

يقارن معامل أكبر من (`>`) بين قيم عددين. إذا كان العدد على اليسار أكبر من العدد على اليمين، يُرجع `true`. وإلا، يُرجع `false`.

مثل معامل المساواة، يقوم معامل أكبر من بتحويل أنواع البيانات للقيم أثناء المقارنة.

**أمثلة**

```js
5   >  3  // true
7   > '3' // true
2   >  3  // false
'1' >  9  // false
```

# --instructions--

أضف معامل أكبر من إلى الأسطر المحددة بحيث تكون جمل الإرجاع منطقية.

# --hints--

يجب أن تُرجع الدالة `testGreaterThan(0)` السلسلة النصية `10 or Under`

```js
assert(testGreaterThan(0) === '10 or Under');
```

يجب أن تُرجع الدالة `testGreaterThan(10)` السلسلة النصية `10 or Under`

```js
assert(testGreaterThan(10) === '10 or Under');
```

يجب أن تُرجع الدالة `testGreaterThan(11)` السلسلة النصية `Over 10`

```js
assert(testGreaterThan(11) === 'Over 10');
```

يجب أن تُرجع الدالة `testGreaterThan(99)` السلسلة النصية `Over 10`

```js
assert(testGreaterThan(99) === 'Over 10');
```

يجب أن تُرجع الدالة `testGreaterThan(100)` السلسلة النصية `Over 10`

```js
assert(testGreaterThan(100) === 'Over 10');
```

يجب أن تُرجع الدالة `testGreaterThan(101)` السلسلة النصية `Over 100`

```js
assert(testGreaterThan(101) === 'Over 100');
```

يجب أن تُرجع الدالة `testGreaterThan(150)` السلسلة النصية `Over 100`

```js
assert(testGreaterThan(150) === 'Over 100');
```

يجب أن تستخدم معامل `>` مرتين على الأقل

```js
assert(__helpers.removeJSComments(code).match(/val\s*>\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testGreaterThan(val) {
  if (val) {  // Change this line
    return "Over 100";
  }

  if (val) {  // Change this line
    return "Over 10";
  }

  return "10 or Under";
}

testGreaterThan(10);
```

# --solutions--

```js
function testGreaterThan(val) {
  if (val > 100) {  // Change this line
    return "Over 100";
  }
  if (val > 10) {  // Change this line
    return "Over 10";
  }
  return "10 or Under";
}
```
