---
id: 56533eb9ac21ba0edf2244bb
title: فراغات الكلمات
challengeType: 1
forumTopicId: 18377
dashedName: word-blanks
---

# --description--

تُزوَّد بجمل تحتوي على بعض الكلمات الناقصة، مثل الأسماء، والأفعال، والصفات، والظروف. ثم تملأ الفراغات بكلمات من اختيارك بحيث تكون الجملة المكتملة منطقية.

انظر إلى هذه الجملة:

```md
It was really ____, and we ____ ourselves ____.
```

تحتوي هذه الجملة على ثلاث كلمات ناقصة - صفة، وفعل، وظرف، ويمكننا إضافة كلمات من اختيارنا لإكمالها. ثم يمكننا تعيين الجملة المكتملة إلى متغير كما يلي:

```js
const sentence = "It was really " + "hot" + ", and we " + "laughed" + " ourselves " + "silly" + ".";
```

# --instructions--

في هذا التحدٍّ، نزودك باسم، وفعل، وصفة، وظرف. تحتاج إلى تكوين جملة كاملة باستخدام كلمات من اختيارك، إلى جانب الكلمات التي نزودك بها.

ستحتاج إلى استخدام معامل التسلسل `+` لبناء سلسلة نصية جديدة، باستخدام المتغيرات المقدمة: `myNoun`، `myAdjective`، `myVerb`، و`myAdverb`. ثم تعيّن السلسلة المشكلة إلى المتغير `wordBlanks`. لا يجب تغيير الكلمات المُعيَّنة للمتغيرات.

ستحتاج أيضًا إلى مراعاة الفراغات في السلسلة النصية، بحيث تحتوي الجملة النهائية على فراغات بين جميع الكلمات. يجب أن تكون النتيجة جملة كاملة.

# --before-each--

```js
const removeAssignments = str => str
  .replace(/myNoun\s*=\s*["']dog["']/g, '')
  .replace(/myAdjective\s*=\s*["']big["']/g, '')
  .replace(/myVerb\s*=\s*["']ran["']/g, '')
  .replace(/myAdverb\s*=\s*["']quickly["']/g, '');
```

# --hints--

يجب أن يكون `wordBlanks` سلسلة نصية.

```js
assert(typeof wordBlanks === 'string');
```

لا يجب تغيير القيم المُعيَّنة لـ `myNoun`، أو `myVerb`، أو `myAdjective`، أو `myAdverb`.

```js
assert(
  myNoun === 'dog' &&
    myVerb === 'ran' &&
    myAdjective === 'big' &&
    myAdverb === 'quickly'
);
```

لا يجب استخدام القيم `dog`، `ran`، `big`، أو `quickly` مباشرةً لإنشاء `wordBlanks`.

```js
const newCode = removeAssignments(__helpers.removeJSComments(code));
assert(
  !/dog/.test(newCode) &&
    !/ran/.test(newCode) &&
    !/big/.test(newCode) &&
    !/quickly/.test(newCode)
);
```

يجب أن يحتوي `wordBlanks` على جميع الكلمات المُعيَّنة للمتغيرات `myNoun`، و`myVerb`، و`myAdjective`، و`myAdverb` مفصولة بحروف غير كلمات (وأي كلمات إضافية من اختيارك).

```js
assert(
  /\bdog\b/.test(wordBlanks) &&
    /\bbig\b/.test(wordBlanks) &&
    /\bran\b/.test(wordBlanks) &&
    /\bquickly\b/.test(wordBlanks)
);
```

# --seed--

## --seed-contents--

```js
const myNoun = "dog";
const myAdjective = "big";
const myVerb = "ran";
const myAdverb = "quickly";

// Only change code below this line
const wordBlanks = ""; // Change this line
// Only change code above this line
```

# --solutions--

```js
const myNoun = "dog";
const myAdjective = "big";
const myVerb = "ran";
const myAdverb = "quickly";

let wordBlanks = "Once there was a " + myNoun + " which was very " + myAdjective + ". ";
wordBlanks += "It " + myVerb + " " + myAdverb + " around the yard.";
```
