---
id: 56533eb9ac21ba0edf2244b0
title: تعيين مركب مع الطرح المعزز
challengeType: 1
forumTopicId: 16660
dashedName: compound-assignment-with-augmented-subtraction
---

# --description--

مثل العامل `+=`، يقوم العامل `-=` بطرح رقم من متغير.

```js
myVar = myVar - 5;
```

سيطرح `5` من `myVar`. يمكن إعادة كتابة ذلك كالتالي:

```js
myVar -= 5;
```

# --instructions--

حوّل التعيينات لـ `a` و `b` و `c` لاستخدام العامل `-=`.

# --hints--

يجب أن تساوي `a` القيمة `5`.

```js
assert(a === 5);
```

يجب أن تساوي `b` القيمة `-6`.

```js
assert(b === -6);
```

يجب أن تساوي `c` القيمة `2`.

```js
assert(c === 2);
```

يجب استخدام العامل `-=` لكل متغير.

```js
assert(__helpers.removeJSComments(code).match(/-=/g).length === 3);
```

لا تعدل الشفرة فوق التعليق المحدد.

```js
assert(
  /let a = 11;/.test(__helpers.removeJSComments(code)) && /let b = 9;/.test(__helpers.removeJSComments(code)) && /let c = 3;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 11;
let b = 9;
let c = 3;

// Only change code below this line
a = a - 6;
b = b - 15;
c = c - 1;
```

# --solutions--

```js
let a = 11;
let b = 9;
let c = 3;

a -= 6;
b -= 15;
c -= 1;
```
