---
id: 56533eb9ac21ba0edf2244b1
title: تعيين مركب مع الضرب المعزز
challengeType: 1
forumTopicId: 16662
dashedName: compound-assignment-with-augmented-multiplication
---

# --description--

المشغل `*=` يضرب متغيرًا بعدد.

```js
myVar = myVar * 5;
```

سيضرب `myVar` في `5`. يمكن إعادة كتابته كالتالي:

```js
myVar *= 5;
```

# --instructions--

حوّل التعيينات لـ `a` و `b` و `c` لاستخدام المشغل `*=`.

# --hints--

يجب أن يساوي `a` القيمة `25`.

```js
assert(a === 25);
```

يجب أن يساوي `b` القيمة `36`.

```js
assert(b === 36);
```

يجب أن يساوي `c` القيمة `46`.

```js
assert(c === 46);
```

يجب استخدام المشغل `*=` لكل متغير.

```js
assert(__helpers.removeJSComments(code).match(/\*=/g).length === 3);
```

لا تعدل الشفرة فوق التعليق المحدد.

```js
assert(
  /let a = 5;/.test(__helpers.removeJSComments(code)) &&
    /let b = 12;/.test(__helpers.removeJSComments(code)) &&
    /let c = 4\.6;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 5;
let b = 12;
let c = 4.6;

// Only change code below this line
a = a * 5;
b = 3 * b;
c = c * 10;
```

# --solutions--

```js
let a = 5;
let b = 12;
let c = 4.6;

a *= 5;
b *= 3;
c *= 10;
```
