---
id: 56533eb9ac21ba0edf2244c2
title: إرجاع قيمة من دالة باستخدام return
challengeType: 1
forumTopicId: 18271
dashedName: return-a-value-from-a-function-with-return
---

# --description--

يمكننا تمرير القيم إلى دالة باستخدام <dfn>المعلمات</dfn>. يمكنك استخدام جملة `return` لإرسال قيمة مرة أخرى من الدالة.

**مثال**

```js
function plusThree(num) {
  return num + 3;
}

const answer = plusThree(5);
```

لدى `answer` القيمة `8`.

تأخذ `plusThree` <dfn>معلمة</dfn> لـ `num` وتُرجع قيمة تساوي `num + 3`.

# --instructions--

أنشئ دالة `timesFive` تقبل معلمة واحدة، تضربها في `5`، وتُرجع القيمة الجديدة.

# --hints--

يجب أن تكون `timesFive` دالة

```js
assert(typeof timesFive === 'function');
```

يجب أن تُرجع `timesFive(5)` القيمة `25`

```js
assert(timesFive(5) === 25);
```

يجب أن تُرجع `timesFive(2)` القيمة `10`

```js
assert(timesFive(2) === 10);
```

يجب أن تُرجع `timesFive(0)` القيمة `0`

```js
assert(timesFive(0) === 0);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function timesFive(num) {
  return num * 5;
}
timesFive(10);
```
