---
id: 56533eb9ac21ba0edf2244dc
title: تسلسل عبارات If Else
challengeType: 1
forumTopicId: 16772
dashedName: chaining-if-else-statements
---

# --description--

يمكن تسلسل عبارات `if/else` معًا لتحقيق منطق معقد. إليك <dfn>رمز زائف</dfn> لعدة عبارات `if` / `else if` متسلسلة:

```js
if (condition1) {
  statement1
} else if (condition2) {
  statement2
} else if (condition3) {
  statement3
. . .
} else {
  statementN
}
```

# --instructions--

اكتب عبارات `if`/`else if` متسلسلة لتلبية الشروط التالية:

`num < 5` - أرجع `Tiny`  
`num < 10` - أرجع `Small`  
`num < 15` - أرجع `Medium`  
`num < 20` - أرجع `Large`  
`num >= 20` - أرجع `Huge`

# --hints--

يجب أن تحتوي على أربعة عبارات `else` على الأقل

```js
assert(__helpers.removeJSComments(code).match(/else/g).length > 3);
```

يجب أن تحتوي على أربعة عبارات `if` على الأقل

```js
assert(__helpers.removeJSComments(code).match(/if/g).length > 3);
```

يجب أن تحتوي على عبارة `return` واحدة على الأقل

```js
assert(__helpers.removeJSComments(code).match(/return/g).length >= 1);
```

يجب أن تُرجع `testSize(0)` السلسلة النصية `Tiny`

```js
assert(testSize(0) === 'Tiny');
```

يجب أن تُرجع `testSize(4)` السلسلة النصية `Tiny`

```js
assert(testSize(4) === 'Tiny');
```

يجب أن تُرجع `testSize(5)` السلسلة النصية `Small`

```js
assert(testSize(5) === 'Small');
```

يجب أن تُرجع `testSize(8)` السلسلة النصية `Small`

```js
assert(testSize(8) === 'Small');
```

يجب أن تُرجع `testSize(10)` السلسلة النصية `Medium`

```js
assert(testSize(10) === 'Medium');
```

يجب أن تُرجع `testSize(14)` السلسلة النصية `Medium`

```js
assert(testSize(14) === 'Medium');
```

يجب أن تُرجع `testSize(15)` السلسلة النصية `Large`

```js
assert(testSize(15) === 'Large');
```

يجب أن تُرجع `testSize(17)` السلسلة النصية `Large`

```js
assert(testSize(17) === 'Large');
```

يجب أن تُرجع `testSize(20)` السلسلة النصية `Huge`

```js
assert(testSize(20) === 'Huge');
```

يجب أن تُرجع `testSize(25)` السلسلة النصية `Huge`

```js
assert(testSize(25) === 'Huge');
```

# --seed--

## --seed-contents--

```js
function testSize(num) {
  // Only change code below this line


  return "Change Me";
  // Only change code above this line
}

testSize(7);
```

# --solutions--

```js
function testSize(num) {
  if (num < 5) {
    return "Tiny";
  } else if (num < 10) {
    return "Small";
  } else if (num < 15) {
    return "Medium";
  } else if (num < 20) {
    return "Large";
  } else {
    return "Huge";
  }
}
```
