---
id: 56533eb9ac21ba0edf2244d2
title: المقارنة باستخدام معامل عدم المساواة
challengeType: 1
forumTopicId: 16787
dashedName: comparison-with-the-inequality-operator
---

# --description--

معامل عدم المساواة (`!=`) هو عكس معامل المساواة. عدم المساواة تعني غير متساوي. يعيد معامل عدم المساواة القيمة `false` عندما يعيد معامل المساواة القيمة `true` والعكس صحيح. مثل معامل المساواة، يحول معامل عدم المساواة أنواع البيانات للقيم أثناء المقارنة.

**أمثلة**

```js
1 !=  2    // true
1 != "1"   // false
1 != '1'   // false
1 != true  // false
0 != false // false
```

# --instructions--

أضف معامل عدم المساواة `!=` في جملة `if` بحيث تُعيد الدالة السلسلة النصية `Not Equal` عندما لا يكون `val` مكافئًا لـ `99`.

# --hints--

`testNotEqual(99)` يجب أن تُعيد السلسلة النصية `Equal`

```js
assert(testNotEqual(99) === 'Equal');
```

`testNotEqual("99")` يجب أن تُعيد السلسلة النصية `Equal`

```js
assert(testNotEqual('99') === 'Equal');
```

`testNotEqual(12)` يجب أن تُعيد السلسلة النصية `Not Equal`

```js
assert(testNotEqual(12) === 'Not Equal');
```

`testNotEqual("12")` يجب أن تُعيد السلسلة النصية `Not Equal`

```js
assert(testNotEqual('12') === 'Not Equal');
```

`testNotEqual("bob")` يجب أن تُعيد السلسلة النصية `Not Equal`

```js
assert(testNotEqual('bob') === 'Not Equal');
```

يجب أن تستخدم معامل `!=`

```js
assert(__helpers.removeJSComments(code).match(/(?!!==)!=/));
```

# --seed--

## --seed-contents--

```js
// Setup
function testNotEqual(val) {
  if (val) { // Change this line
    return "Not Equal";
  }
  return "Equal";
}

testNotEqual(10);
```

# --solutions--

```js
function testNotEqual(val) {
  if (val != 99) {
    return "Not Equal";
  }
  return "Equal";
}
```
