---
id: 56533eb9ac21ba0edf2244af
title: تعيين مركب مع الجمع المعزز
challengeType: 1
forumTopicId: 16661
dashedName: compound-assignment-with-augmented-addition
---

# --description--

في البرمجة، من الشائع استخدام التعيينات لتعديل محتويات متغير. تذكر أن كل شيء على يمين علامة المساواة يُحلل أولًا، لذلك يمكننا القول:

```js
myVar = myVar + 5;
```

لإضافة `5` إلى `myVar`. بما أن هذا نمط شائع جدًا، هناك عوامل تقوم بعملية رياضية وتعيين في خطوة واحدة.

أحد هذه العوامل هو العامل `+=`.

```js
let myVar = 1;
myVar += 5;
console.log(myVar);
```

سيُعرض الرقم `6` في وحدة التحكم.

# --instructions--

حوّل التعيينات لـ `a` و `b` و `c` لاستخدام العامل `+=`.

# --hints--

يجب أن يساوي `a` القيمة `15`.

```js
assert(a === 15);
```

يجب أن يساوي `b` القيمة `26`.

```js
assert(b === 26);
```

يجب أن يساوي `c` القيمة `19`.

```js
assert(c === 19);
```

يجب استخدام العامل `+=` لكل متغير.

```js
assert(__helpers.removeJSComments(code).match(/\+=/g).length === 3);
```

لا تعدل الشفرة فوق التعليق المحدد.

```js
assert(
  /let a = 3;/.test(__helpers.removeJSComments(code)) &&
    /let b = 17;/.test(__helpers.removeJSComments(code)) &&
    /let c = 12;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 3;
let b = 17;
let c = 12;

// Only change code below this line
a = a + 12;
b = 9 + b;
c = c + 7;
```

# --solutions--

```js
let a = 3;
let b = 17;
let c = 12;

a += 12;
b += 9;
c += 7;
```
