---
id: 56533eb9ac21ba0edf2244ac
title: زيادة رقم باستخدام JavaScript
challengeType: 1
forumTopicId: 18201
dashedName: increment-a-number-with-javascript
---

# --description--

يمكنك بسهولة <dfn>زيادة</dfn> أو إضافة واحد إلى متغير باستخدام معامل `++`.

```js
i++;
```

يعادل

```js
i = i + 1;
```

**ملاحظة:** تصبح السطر بأكمله `i++;`، مما يلغي الحاجة إلى علامة المساواة.

# --instructions--

غيّر الشفرة لاستخدام معامل `++` على `myVar`.

# --hints--

يجب أن يساوي `myVar` القيمة `88`.

```js
assert(myVar === 88);
```

يجب ألا تستخدم معامل التعيين.

```js
assert(
  /let\s+myVar\s*=\s*87;\s*\/*.*\s*([+]{2}\s*myVar|myVar\s*[+]{2})/.test(__helpers.removeJSComments(code))
);
```

يجب أن تستخدم معامل `++`.

```js
assert(/[+]{2}\s*myVar|myVar\s*[+]{2}/.test(__helpers.removeJSComments(code)));
```

يجب ألا تغيّر الشفرة فوق التعليق المحدد.

```js
assert(/let myVar = 87;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let myVar = 87;

// Only change code below this line
myVar = myVar + 1;
```

# --solutions--

```js
let myVar = 87;
myVar++;
```
