---
id: 56533eb9ac21ba0edf2244d8
title: المقارنات باستخدام عامل الربط المنطقي And
challengeType: 1
forumTopicId: 16799
dashedName: comparisons-with-the-logical-and-operator
---

# --description--

أحيانًا تحتاج إلى اختبار أكثر من شيء في نفس الوقت. عامل الربط المنطقي <dfn>and</dfn> (`&&`) يُرجع `true` فقط إذا كانت <dfn>الوسائط</dfn> على يمينه ويساره صحيحة.

يمكن تحقيق نفس النتيجة عن طريق التسكين لجملة `if` داخل جملة `if` أخرى.

```js
if (num > 5) {
  if (num < 10) {
    return "Yes";
  }
}
return "No";
```

ستُرجع هذه الشفرة `Yes` إذا كان `num` أكبر من `5` وأقل من `10`. يمكن كتابة نفس المنطق باستخدام عامل الربط المنطقي <dfn>and</dfn>.

```js
if (num > 5 && num < 10) {
  return "Yes";
}
return "No";
```

# --instructions--

استبدل جملتي if بواحدة تستخدم عامل `&&`، بحيث تُرجع السلسلة النصية `Yes` إذا كان `val` أقل من أو يساوي `50` وأكبر من أو يساوي `25`. وإلا، تُرجع السلسلة النصية `No`.

# --hints--

يجب استخدام عامل `&&` مرة واحدة

```js
assert(__helpers.removeJSComments(code).match(/&&/g).length === 1);
```

يجب أن تحتوي على جملة `if` واحدة فقط

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`testLogicalAnd(0)` يجب أن تُرجع السلسلة النصية `No`

```js
assert(testLogicalAnd(0) === 'No');
```

`testLogicalAnd(24)` يجب أن تُرجع السلسلة النصية `No`

```js
assert(testLogicalAnd(24) === 'No');
```

`testLogicalAnd(25)` يجب أن تُرجع السلسلة النصية `Yes`

```js
assert(testLogicalAnd(25) === 'Yes');
```

`testLogicalAnd(30)` يجب أن تُرجع السلسلة النصية `Yes`

```js
assert(testLogicalAnd(30) === 'Yes');
```

`testLogicalAnd(50)` يجب أن تُرجع السلسلة النصية `Yes`

```js
assert(testLogicalAnd(50) === 'Yes');
```

`testLogicalAnd(51)` يجب أن تُرجع السلسلة النصية `No`

```js
assert(testLogicalAnd(51) === 'No');
```

`testLogicalAnd(75)` يجب أن تُرجع السلسلة النصية `No`

```js
assert(testLogicalAnd(75) === 'No');
```

`testLogicalAnd(80)` يجب أن تُرجع السلسلة النصية `No`

```js
assert(testLogicalAnd(80) === 'No');
```

# --seed--

## --seed-contents--

```js
function testLogicalAnd(val) {
  // Only change code below this line

  if (val) {
    if (val) {
      return "Yes";
    }
  }

  // Only change code above this line
  return "No";
}

testLogicalAnd(10);
```

# --solutions--

```js
function testLogicalAnd(val) {
  if (val >= 25 && val <= 50) {
    return "Yes";
  }
  return "No";
}
```
