---
id: 56533eb9ac21ba0edf2244cb
title: التلاعب بالكائنات المعقدة
challengeType: 1
forumTopicId: 18208
dashedName: manipulating-complex-objects
---

# --description--

أحيانًا قد تريد تخزين بيانات في <dfn>هيكل بيانات</dfn> مرن. كائن JavaScript هو طريقة واحدة للتعامل مع البيانات المرنة. يسمح بتركيبات عشوائية من <dfn>سلاسل نصية</dfn>، <dfn>أرقام</dfn>، <dfn>قيم منطقية</dfn>، <dfn>مصفوفات</dfn>، <dfn>دوال</dfn>، و<dfn>كائنات</dfn>.

إليك مثالًا على هيكل بيانات معقد:

```js
const ourMusic = [
  {
    "artist": "Daft Punk",
    "title": "Homework",
    "release_year": 1997,
    "formats": [ 
      "CD", 
      "Cassette", 
      "LP"
    ],
    "gold": true
  }
];
```

هذه مصفوفة تحتوي على كائن واحد بداخلها. يحتوي الكائن على عدة قطع من <dfn>بيانات وصفية</dfn> عن ألبوم. كما يحتوي على مصفوفة `formats` مسكنة. إذا أردت إضافة سجلات ألبومات أكثر، يمكنك فعل ذلك بإضافة سجلات إلى المصفوفة على المستوى الأعلى. الكائنات تخزن البيانات في خاصية، والتي لها صيغة مفتاح-قيمة. في المثال أعلاه، `"artist": "Daft Punk"` هي خاصية لها مفتاح `artist` وقيمة `Daft Punk`.

**ملاحظة:** يجب وضع فاصلة بعد كل كائن في المصفوفة، إلا إذا كان هو الكائن الأخير في المصفوفة.

# --instructions--

أضف ألبومًا جديدًا إلى مصفوفة `myMusic`. أضف سلاسل نصية لـ `artist` و `title`، ورقم لـ `release_year`، ومصفوفة من السلاسل النصية لـ `formats`.

# --hints--

يجب أن تكون `myMusic` مصفوفة

```js
assert(Array.isArray(myMusic));
```

يجب أن تحتوي `myMusic` على عنصرين على الأقل

```js
assert(myMusic.length > 1);
```

يجب أن تكون العناصر في مصفوفة `myMusic` كائنات

```js
myMusic.forEach(object => {assert.typeOf(object, 'object')})
```

يجب أن يحتوي كائنك في `myMusic` على 4 خصائص على الأقل

```js
myMusic.forEach(object => {assert(Object.keys(object).length > 3); });
```

يجب أن يحتوي كائنك في `myMusic` على الخاصية `artist` التي هي سلسلة نصية

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['artist']);
  assert.typeOf(object.artist, 'string')
})
```

يجب أن يحتوي كائنك في `myMusic` على الخاصية `title` التي هي سلسلة نصية

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['title']);
  assert.typeOf(object.title, 'string')
})
```

يجب أن يحتوي كائنك في `myMusic` على الخاصية `release_year` التي هي رقم

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['release_year']);
  assert.typeOf(object.release_year, 'number')
})
```

يجب أن يحتوي كائنك في `myMusic` على خاصية `formats` التي هي مصفوفة

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['formats']);
  assert.typeOf(object.formats, 'array')
})
```

يجب أن تكون `formats` مصفوفة من سلاسل نصية تحتوي على عنصرين على الأقل

```js
myMusic.forEach(object => {
  object.formats.forEach(format => {
    assert.typeOf(format, 'string')
  });
  assert.isAtLeast(object.formats.length, 2)
})
```

# --seed--

## --seed-contents--

```js
const myMusic = [
  {
    "artist": "Billy Joel",
    "title": "Piano Man",
    "release_year": 1973,
    "formats": [
      "CD",
      "8T",
      "LP"
    ],
    "gold": true
  }
];
```

# --solutions--

```js
const myMusic = [
  {
    "artist": "Billy Joel",
    "title": "Piano Man",
    "release_year": 1973,
    "formats": [
      "CS",
      "8T",
      "LP" ],
    "gold": true
  },
  {
    "artist": "ABBA",
    "title": "Ring Ring",
    "release_year": 1973,
    "formats": [
      "CS",
      "8T",
      "LP",
    "CD",
  ]
  }
];
```
