---
id: 56533eb9ac21ba0edf2244de
title: إضافة خيار افتراضي في جمل التبديل
challengeType: 1
forumTopicId: 16653
dashedName: adding-a-default-option-in-switch-statements
---

# --description--

في جملة `switch` قد لا تتمكن من تحديد كل القيم الممكنة كجمل `case`. بدلًا من ذلك، يمكنك إضافة جملة `default` التي تُنفذ إذا لم تُوجد أي جمل `case` مطابقة. فكر فيها كأنها جملة `else` النهائية في سلسلة `if/else`.

يجب أن تكون جملة `default` هي الحالة الأخيرة.

```js
switch (num) {
  case value1:
    statement1;
    break;
  case value2:
    statement2;
    break;
...
  default:
    defaultStatement;
    break;
}
```

# --instructions--

اكتب جملة switch لتعيين `answer` للحالات التالية:  
`a` - `apple`  
`b` - `bird`  
`c` - `cat`  
`default` - `stuff`

# --hints--

يجب أن تُرجع الدالة `switchOfStuff("a")` السلسلة النصية `apple`

```js
assert(switchOfStuff('a') === 'apple');
```

يجب أن تُرجع الدالة `switchOfStuff("b")` السلسلة النصية `bird`

```js
assert(switchOfStuff('b') === 'bird');
```

يجب أن تُرجع الدالة `switchOfStuff("c")` السلسلة النصية `cat`

```js
assert(switchOfStuff('c') === 'cat');
```

يجب أن تُرجع الدالة `switchOfStuff("d")` السلسلة النصية `stuff`

```js
assert(switchOfStuff('d') === 'stuff');
```

يجب أن تُرجع الدالة `switchOfStuff(4)` السلسلة النصية `stuff`

```js
assert(switchOfStuff(4) === 'stuff');
```

لا يجب أن تستخدم أي جمل `if` أو `else`

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

يجب أن تستخدم جملة `default`

```js
assert(switchOfStuff('string-to-trigger-default-case') === 'stuff');
```

يجب أن تحتوي على 3 جمل `break` على الأقل

```js
assert(__helpers.removeJSComments(code).match(/break/g).length > 2);
```

# --seed--

## --seed-contents--

```js
function switchOfStuff(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

switchOfStuff(1);
```

# --solutions--

```js
function switchOfStuff(val) {
  let answer = "";

  switch(val) {
    case "a":
      answer = "apple";
      break;
    case "b":
      answer = "bird";
      break;
    case "c":
      answer = "cat";
      break;
    default:
      answer = "stuff";
  }
  return answer;
}
```
