---
id: 56533eb9ac21ba0edf2244ca
title: استخدام الكائنات للبحث
challengeType: 1
forumTopicId: 18373
dashedName: using-objects-for-lookups
---

# --description--

يمكن اعتبار الكائنات كمخزن مفتاح/قيمة، مثل القاموس. إذا كان لديك بيانات جدولية، يمكنك استخدام كائن للبحث عن القيم بدلاً من جملة `switch` أو سلسلة `if/else`. هذا يكون مفيدًا جدًا عندما تعرف أن بيانات المُدخَل محدودة لنطاق معين.

إليك مثال على كائن مقال:

```js
const article = {
  "title": "How to create objects in JavaScript",
  "link": "https://www.freecodecamp.org/news/a-complete-guide-to-creating-objects-in-javascript-b0e2450655e8/",
  "author": "Kaashan Hussain",
  "language": "JavaScript",
  "tags": "TECHNOLOGY",
  "createdAt": "NOVEMBER 28, 2018"
};

const articleAuthor = article["author"];
const articleLink = article["link"];

const value = "title";
const valueLookup = article[value];
```

`articleAuthor` هي السلسلة النصية `Kaashan Hussain`، و`articleLink` هي السلسلة النصية `https://www.freecodecamp.org/news/a-complete-guide-to-creating-objects-in-javascript-b0e2450655e8/`، و`valueLookup` هي السلسلة النصية `How to create objects in JavaScript`.

# --instructions--

حوّل جملة التبديل إلى كائن يسمى `lookup`. استخدمه للبحث عن `val` وعيّن السلسلة النصية المرتبطة إلى المتغير `result`.

# --hints--

`phoneticLookup("alpha")` يجب أن تساوي السلسلة النصية `Adams`

```js
assert(phoneticLookup('alpha') === 'Adams');
```

`phoneticLookup("bravo")` يجب أن تساوي السلسلة النصية `Boston`

```js
assert(phoneticLookup('bravo') === 'Boston');
```

`phoneticLookup("charlie")` يجب أن تساوي السلسلة النصية `Chicago`

```js
assert(phoneticLookup('charlie') === 'Chicago');
```

`phoneticLookup("delta")` يجب أن تساوي السلسلة النصية `Denver`

```js
assert(phoneticLookup('delta') === 'Denver');
```

`phoneticLookup("echo")` يجب أن تساوي السلسلة النصية `Easy`

```js
assert(phoneticLookup('echo') === 'Easy');
```

`phoneticLookup("foxtrot")` يجب أن تساوي السلسلة النصية `Frank`

```js
assert(phoneticLookup('foxtrot') === 'Frank');
```

`phoneticLookup("")` يجب أن تساوي `undefined`

```js
assert(typeof phoneticLookup('') === 'undefined');
```

لا يجب تعديل جملة `return`

```js
assert(__helpers.removeJSComments(code).match(/return\sresult;/));
```

لا يجب استخدام جمل `case` أو `switch` أو `if`

```js
assert(
  !/case|switch|if/g.test(__helpers.removeJSComments(code).replace(/([/]{2}.*)|([/][*][^/*]*[*][/])/g, ''))
);
```

# --seed--

## --seed-contents--

```js
// Setup
function phoneticLookup(val) {
  let result = "";

  // Only change code below this line
  switch(val) {
    case "alpha":
      result = "Adams";
      break;
    case "bravo":
      result = "Boston";
      break;
    case "charlie":
      result = "Chicago";
      break;
    case "delta":
      result = "Denver";
      break;
    case "echo":
      result = "Easy";
      break;
    case "foxtrot":
      result = "Frank";
  }

  // Only change code above this line
  return result;
}

phoneticLookup("charlie");
```

# --solutions--

```js
function phoneticLookup(val) {
  let result = "";

  const lookup = {
    alpha: "Adams",
    bravo: "Boston",
    charlie: "Chicago",
    delta: "Denver",
    echo: "Easy",
    foxtrot: "Frank"
  };

  result = lookup[val];

  return result;
}
```
