---
id: 56533eb9ac21ba0edf2244c6
title: الوقوف في الصف
challengeType: 1
forumTopicId: 18307
dashedName: stand-in-line
---

# --description--

في علوم الحاسوب، <dfn>الطابور</dfn> هو <dfn>هيكل بيانات</dfn> مجرد تُحفظ فيه العناصر بترتيب معين. يمكن إضافة عناصر جديدة في نهاية الطابور وإزالة العناصر القديمة من بداية الطابور.

# --instructions--

اكتب دالة `nextInLine` تأخذ مصفوفة (`arr`) ورقم (`item`) كمعلمات.

أضف الرقم إلى نهاية المصفوفة، ثم أزل العنصر الأول من المصفوفة.

يجب أن تُرجع دالة `nextInLine` العنصر الذي أُزيل.

# --hints--

`nextInLine([], 5)` يجب أن تُرجع رقمًا.

```js
assert.isNumber(nextInLine([], 5));
```

`nextInLine([], 1)` يجب أن تُرجع `1`

```js
assert(nextInLine([], 1) === 1);
```

`nextInLine([2], 1)` يجب أن تُرجع `2`

```js
assert(nextInLine([2], 1) === 2);
```

`nextInLine([5,6,7,8,9], 1)` يجب أن تُرجع `5`

```js
assert(nextInLine([5, 6, 7, 8, 9], 1) === 5);
```

# --seed--

## --seed-contents--

```js
function nextInLine(arr, item) {
  // Only change code below this line
  
  return item;
  // Only change code above this line
}

// Setup
let testArr = [1, 2, 3, 4, 5];

// Display code
console.log("Before: " + JSON.stringify(testArr));
console.log(nextInLine(testArr, 6));
console.log("After: " + JSON.stringify(testArr));
```

# --solutions--

```js
let testArr = [1, 2, 3, 4, 5];

function nextInLine(arr, item) {
    arr.push(item);
    return arr.shift();
}
```
