---
id: 56533eb9ac21ba0edf2244cc
title: الوصول إلى الكائنات المسكنة
challengeType: 1
forumTopicId: 16161
dashedName: accessing-nested-objects
---

# --description--

يمكن الوصول إلى الخصائص الفرعية للكائنات عبر تسلسل ترميز النقطة أو ترميز الأقواس.

إليك كائن مسكن:

```js
const ourStorage = {
  "desk": {
    "drawer": "stapler"
  },
  "cabinet": {
    "top drawer": { 
      "folder1": "a file",
      "folder2": "secrets"
    },
    "bottom drawer": "soda"
  }
};

ourStorage.cabinet["top drawer"].folder2;
ourStorage.desk.drawer;
```

`ourStorage.cabinet["top drawer"].folder2` ستكون السلسلة النصية `secrets`، و`ourStorage.desk.drawer` ستكون السلسلة النصية `stapler`.

# --instructions--

ادخل إلى الكائن `myStorage` وعيّن محتويات الخاصية `glove box` إلى المتغير `gloveBoxContents`. استخدم ترميز النقطة لكل الخصائص حيثما أمكن، وإلا فاستخدم ترميز الأقواس.

# --hints--

يجب أن تساوي قيمة `gloveBoxContents` السلسلة النصية `maps`.

```js
assert(gloveBoxContents === 'maps');
```

يجب أن يستخدم شفرتك ترميز النقطة، حيثما أمكن، للوصول إلى `myStorage`.

```js
assert.match(code, /myStorage\.car\.inside/);
```

يجب أن يُعلن عن `gloveBoxContents` باستخدام `const` كما هو.

```js
assert.match(code, /const\s+gloveBoxContents\s*=/);
```

لا يجب أن تغيّر الكائن `myStorage`.

```js
const expectedMyStorage = {
  "car":{
    "inside":{
      "glove box":"maps",
      "passenger seat":"crumbs"
    },
    "outside":{
      "trunk":"jack"
    }
  }
};
assert.deepStrictEqual(myStorage, expectedMyStorage);
```

# --seed--

## --seed-contents--

```js
const myStorage = {
  "car": {
    "inside": {
      "glove box": "maps",
      "passenger seat": "crumbs"
     },
    "outside": {
      "trunk": "jack"
    }
  }
};

const gloveBoxContents = undefined;
```

# --solutions--

```js
const myStorage = {
  "car":{
    "inside":{
      "glove box":"maps",
      "passenger seat":"crumbs"
    },
    "outside":{
      "trunk":"jack"
    }
  }
};
const gloveBoxContents = myStorage.car.inside["glove box"];
```
