---
id: 56533eb9ac21ba0edf2244b5
title: الهروب من علامات الاقتباس الحرفية في السلاسل النصية
challengeType: 1
forumTopicId: 17568
dashedName: escaping-literal-quotes-in-strings
---

# --description--

عند تعريف سلسلة نصية يجب أن تبدأ وتنتهي بعلامة اقتباس مفردة أو مزدوجة. ماذا يحدث عندما تحتاج إلى علامة اقتباس حرفية: `"` أو `'` داخل السلسلة النصية؟

في JavaScript، يمكنك <dfn>الهروب</dfn> من علامة الاقتباس حتى لا تُعتبر نهاية السلسلة النصية بوضع <dfn>شرطة مائلة عكسية</dfn> (`\`) أمام علامة الاقتباس.

```js
const sampleStr = "Alan said, \"Peter is learning JavaScript\".";
```

هذا يشير إلى JavaScript أن علامة الاقتباس التالية ليست نهاية السلسلة، بل يجب أن تظهر داخل السلسلة. لذا إذا طبعت هذا في وحدة التحكم، ستحصل على:

```js
Alan said, "Peter is learning JavaScript".
```

# --instructions--

استخدم <dfn>شرطات مائلة عكسية</dfn> لتعيين سلسلة نصية للمتغير `myStr` بحيث إذا طبعتها في وحدة التحكم، سترى:

```js
I am a "double quoted" string inside "double quotes".
```

# --hints--

يجب أن تستخدم علامتي اقتباس مزدوجتين (`"`) وأربع علامات اقتباس مزدوجة مهروبة (`\"`).

```js
assert(__helpers.removeJSComments(code).match(/\\"/g).length === 4 && __helpers.removeJSComments(code).match(/[^\\]"/g).length === 2);
```

يجب أن يحتوي المتغير `myStr` على السلسلة النصية: `I am a "double quoted" string inside "double quotes".`

```js
assert(/I am a "double quoted" string inside "double quotes(\."|"\.)$/.test(myStr));
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "I am a \"double quoted\" string inside \"double quotes\".";
```
