---
id: 56533eb9ac21ba0edf2244bf
title: النطاق المحلي والدوال
challengeType: 1
forumTopicId: 18227
dashedName: local-scope-and-functions
---

# --description--

المتغيرات التي تُعلن داخل دالة، بالإضافة إلى معلمات الدالة، لها <dfn>نطاق</dfn> محلي. هذا يعني أنها تظهر فقط داخل تلك الدالة.

إليك دالة `myTest` تحتوي على متغير محلي اسمه `loc`.

```js
function myTest() {
  const loc = "foo";
  console.log(loc);
}

myTest();
console.log(loc);
```

استدعاء الدالة `myTest()` سيعرض السلسلة النصية `foo` في وحدة التحكم. سطر `console.log(loc)` (خارج دالة `myTest`) سيرمي خطأً، لأن `loc` غير معرف خارج الدالة.

# --instructions--

المحرر يحتوي على اثنين من `console.log` لمساعدتك على رؤية ما يحدث. راقب وحدة التحكم أثناء كتابة الشفرة لترى كيف تتغير. أعلن متغيرًا محليًا `myVar` داخل `myLocalScope` وشغّل الاختبارات.

**ملاحظة:** ستظل وحدة التحكم تعرض `ReferenceError: myVar is not defined`، لكن هذا لن يسبب فشل الاختبارات.

# --hints--

يجب ألا تحتوي الشفرة على متغير `myVar` عام.

```js
function declared() {
  myVar;
}

assert.throws(declared, ReferenceError);
```

يجب إضافة متغير محلي `myVar`.

```js
assert(
  /functionmyLocalScope\(\)\{.*(var|let|const)myVar[\s\S]*}/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

# --seed--

## --seed-contents--

```js
function myLocalScope() {
  // Only change code below this line

  console.log('inside myLocalScope', myVar);
}
myLocalScope();

// Run and check the console
// myVar is not defined outside of myLocalScope
console.log('outside myLocalScope', myVar);
```

# --solutions--

```js
function myLocalScope() {
  // Only change code below this line
  let myVar;
  console.log('inside myLocalScope', myVar);
}
myLocalScope();

// Run and check the console
// myVar is not defined outside of myLocalScope
console.log('outside myLocalScope', myVar);
```
