---
id: 56533eb9ac21ba0edf2244df
title: خيارات متطابقة متعددة في جمل التبديل
challengeType: 1
forumTopicId: 18242
dashedName: multiple-identical-options-in-switch-statements
---

# --description--

إذا حُذفت جملة `break` من `case` في جملة `switch`، تُنفذ جملة (أو جمل) `case` التالية حتى تُصادف جملة `break`. إذا كان لديك مُدخلات متعددة لها نفس الناتج، يمكنك تمثيلها في جملة `switch` بهذه الطريقة:

```js
let result = "";
switch (val) {
  case 1:
  case 2:
  case 3:
    result = "1, 2, or 3";
    break;
  case 4:
    result = "4 alone";
}
```

الحالات 1 و 2 و 3 ستنتج جميعها نفس النتيجة.

# --instructions--

اكتب جملة التبديل لتعيين `answer` للنطاقات التالية:  
`1-3` - `Low`  
`4-6` - `Mid`  
`7-9` - `High`

**ملاحظة:** ستحتاج إلى وجود جملة `case` لكل رقم في النطاق.

# --hints--

`sequentialSizes(1)` يجب أن تُرجع السلسلة النصية `Low`

```js
assert(sequentialSizes(1) === 'Low');
```

`sequentialSizes(2)` يجب أن تُرجع السلسلة النصية `Low`

```js
assert(sequentialSizes(2) === 'Low');
```

`sequentialSizes(3)` يجب أن تُرجع السلسلة النصية `Low`

```js
assert(sequentialSizes(3) === 'Low');
```

`sequentialSizes(4)` يجب أن تُرجع السلسلة النصية `Mid`

```js
assert(sequentialSizes(4) === 'Mid');
```

`sequentialSizes(5)` يجب أن تُرجع السلسلة النصية `Mid`

```js
assert(sequentialSizes(5) === 'Mid');
```

`sequentialSizes(6)` يجب أن تُرجع السلسلة النصية `Mid`

```js
assert(sequentialSizes(6) === 'Mid');
```

`sequentialSizes(7)` يجب أن تُرجع السلسلة النصية `High`

```js
assert(sequentialSizes(7) === 'High');
```

`sequentialSizes(8)` يجب أن تُرجع السلسلة النصية `High`

```js
assert(sequentialSizes(8) === 'High');
```

`sequentialSizes(9)` يجب أن تُرجع السلسلة النصية `High`

```js
assert(sequentialSizes(9) === 'High');
```

يجب ألا تستخدم أي جمل `if` أو `else`

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

يجب أن تحتوي على تسع جمل `case`

```js
assert(__helpers.removeJSComments(code).match(/case/g).length === 9);
```

# --seed--

## --seed-contents--

```js
function sequentialSizes(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

sequentialSizes(1);
```

# --solutions--

```js
function sequentialSizes(val) {
  let answer = "";

  switch (val) {
    case 1:
    case 2:
    case 3:
      answer = "Low";
      break;
    case 4:
    case 5:
    case 6:
      answer = "Mid";
      break;
    case 7:
    case 8:
    case 9:
      answer = "High";
  }

  return answer;
}
```
