---
id: 56533eb9ac21ba0edf2244b8
title: دمج السلاسل النصية باستخدام عامل الجمع والتعيين
challengeType: 1
forumTopicId: 16803
dashedName: concatenating-strings-with-the-plus-equals-operator
---

# --description--

يمكنك أيضًا استخدام العامل `+=` لـ <dfn>دمج</dfn> سلسلة نصية في نهاية متغير سلسلة نصية موجود. هذا مفيد جدًا لتقسيم سلسلة طويلة على عدة أسطر.

**ملاحظة:** انتبه للمسافات. الدمج لا يضيف مسافات بين السلاسل المدمجة، لذا عليك إضافتها بنفسك.

مثال:

```js
let ourStr = "I come first. ";
ourStr += "I come second.";
```

الآن يحتوي `ourStr` على قيمة السلسلة النصية `I come first. I come second.`.

# --instructions--

ابنِ `myStr` على عدة أسطر بدمج هاتين السلسلتين النصيتين: `This is the first sentence.` و `This is the second sentence.` باستخدام العامل `+=`. استخدم العامل `+=` كما هو موضح في المثال وتأكد من إضافة مسافة بين السلسلتين. ابدأ بتعيين السلسلة الأولى إلى `myStr`، ثم أضف السلسلة الثانية.

# --hints--

يجب أن يحتوي `myStr` على حرف مسافة واحد بين السلسلتين.

```js
assert(/sentence\. This/.test(myStr));
```

يجب أن يحتوي `myStr` على قيمة السلسلة النصية `This is the first sentence. This is the second sentence.`

```js
assert(myStr === 'This is the first sentence. This is the second sentence.');
```

يجب أن تستخدم العامل `+=` لبناء `myStr`.

```js
assert(__helpers.removeJSComments(code).match(/myStr\s*\+=\s*(["']).*\1/g));
```

# --seed--

## --seed-contents--

```js
let myStr;
```

# --solutions--

```js
let myStr = "This is the first sentence. ";
myStr += "This is the second sentence.";
```
