---
id: 56533eb9ac21ba0edf2244ed
title: إلحاق المتغيرات بالسلاسل النصية
challengeType: 1
forumTopicId: 16656
dashedName: appending-variables-to-strings
---

# --description--

تمامًا كما يمكننا بناء سلسلة نصية عبر عدة أسطر من <dfn>ثوابت</dfn> السلاسل النصية، يمكننا أيضًا إلحاق المتغيرات بسلسلة نصية باستخدام معامل الجمع واليساوي (`+=`).

مثال:

```js
const anAdjective = "awesome!";
let ourStr = "freeCodeCamp is ";
ourStr += anAdjective;
```

سيكون لـ `ourStr` القيمة `freeCodeCamp is awesome!`.

# --instructions--

عيّن `someAdjective` إلى سلسلة نصية لا تقل عن 3 أحرف وألحقها بـ `myStr` باستخدام معامل `+=`.

# --hints--

يجب أن يكون `someAdjective` سلسلة نصية لا تقل عن 3 أحرف.

```js
assert(typeof someAdjective !== 'undefined' && someAdjective.length > 2);
```

يجب إلحاق `someAdjective` بـ `myStr` باستخدام معامل `+=`.

```js
assert(__helpers.removeJSComments(code).match(/myStr\s*\+=\s*someAdjective\s*/).length > 0);
```

# --seed--

## --seed-contents--

```js
// Change code below this line
const someAdjective = "";
let myStr = "Learning to code is ";
```

# --solutions--

```js
const someAdjective = "neat";
let myStr = "Learning to code is ";
myStr += someAdjective;
```
