---
id: 56533eb9ac21ba0edf2244b9
title: بناء سلاسل نصية باستخدام المتغيرات
challengeType: 1
forumTopicId: 16805
dashedName: constructing-strings-with-variables
---

# --description--

أحيانًا تحتاج إلى بناء سلسلة نصية. باستخدام معامل التسلسل (`+`)، يمكنك إدخال متغير واحد أو أكثر داخل السلسلة النصية التي تبنيها.

مثال:

```js
const ourName = "freeCodeCamp";
const ourStr = "Hello, our name is " + ourName + ", how are you?";
```

سيكون لـ `ourStr` قيمة السلسلة النصية `Hello, our name is freeCodeCamp, how are you?`.

# --instructions--

عيّن `myName` إلى سلسلة نصية تساوي اسمك وابنِ `myStr` بحيث يكون `myName` بين السلسلتين النصيتين `My name is ` و ` and I am well!`

# --hints--

يجب أن يكون `myName` سلسلة نصية لا تقل عن 3 أحرف.

```js
assert(typeof myName !== 'undefined' && myName.length > 2);
```

يجب أن تستخدم معاملَي `+` لبناء `myStr` بحيث يحتوي على `myName` بداخله.

```js
assert(__helpers.removeJSComments(code).match(/["']\s*\+\s*myName\s*\+\s*["']/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const myName = "";
const myStr = "";
```

# --solutions--

```js
const myName = "Bob";
const myStr = "My name is " + myName + " and I am well!";
```
