---
id: 56533eb9ac21ba0edf2244e0
title: استبدال سلاسل If Else بـ Switch
challengeType: 1
forumTopicId: 18266
dashedName: replacing-if-else-chains-with-switch
---

# --description--

إذا كان لديك العديد من الخيارات للاختيار من بينها، يمكن أن يكون من الأسهل كتابة عبارة `switch` بدلاً من العديد من عبارات `if`/`else if` المتسلسلة. التالي:

```js
if (val === 1) {
  answer = "a";
} else if (val === 2) {
  answer = "b";
} else {
  answer = "c";
}
```

يمكن استبداله بـ:

```js
switch (val) {
  case 1:
    answer = "a";
    break;
  case 2:
    answer = "b";
    break;
  default:
    answer = "c";
}
```

# --instructions--

غيّر عبارات `if`/`else if` المتسلسلة إلى عبارة `switch`.

# --hints--

يجب ألا تستخدم أي عبارات `else` في أي مكان في المحرر

```js
assert(!/else/g.test(__helpers.removeJSComments(code)));
```

يجب ألا تستخدم أي عبارات `if` في أي مكان في المحرر

```js
assert(!/if/g.test(__helpers.removeJSComments(code)));
```

يجب أن تحتوي على أربعة عبارات `break` على الأقل

```js
assert(__helpers.removeJSComments(code).match(/break/g).length >= 4);
```

`chainToSwitch("bob")` يجب أن تُرجع السلسلة النصية `Marley`

```js
assert(chainToSwitch('bob') === 'Marley');
```

`chainToSwitch(42)` يجب أن تُرجع السلسلة النصية `The Answer`

```js
assert(chainToSwitch(42) === 'The Answer');
```

`chainToSwitch(1)` يجب أن تُرجع السلسلة النصية `There is no #1`

```js
assert(chainToSwitch(1) === 'There is no #1');
```

`chainToSwitch(99)` يجب أن تُرجع السلسلة النصية `Missed me by this much!`

```js
assert(chainToSwitch(99) === 'Missed me by this much!');
```

`chainToSwitch(7)` يجب أن تُرجع السلسلة النصية `Ate Nine`

```js
assert(chainToSwitch(7) === 'Ate Nine');
```

`chainToSwitch("John")` يجب أن تُرجع `""` (سلسلة نصية فارغة)

```js
assert(chainToSwitch('John') === '');
```

`chainToSwitch(156)` يجب أن تُرجع `""` (سلسلة نصية فارغة)

```js
assert(chainToSwitch(156) === '');
```

# --seed--

## --seed-contents--

```js
function chainToSwitch(val) {
  let answer = "";
  // Only change code below this line

  if (val === "bob") {
    answer = "Marley";
  } else if (val === 42) {
    answer = "The Answer";
  } else if (val === 1) {
    answer = "There is no #1";
  } else if (val === 99) {
    answer = "Missed me by this much!";
  } else if (val === 7) {
    answer = "Ate Nine";
  }

  // Only change code above this line
  return answer;
}

chainToSwitch(7);
```

# --solutions--

```js
function chainToSwitch(val) {
  let answer = "";

  switch (val) {
    case "bob":
      answer = "Marley";
      break;
    case 42:
      answer = "The Answer";
      break;
    case 1:
      answer = "There is no #1";
      break;
    case 99:
      answer = "Missed me by this much!";
      break;
    case 7:
      answer = "Ate Nine";
  }
  return answer;
}
```
