---
id: 56533eb9ac21ba0edf2244a9
title: تهيئة المتغيرات باستخدام معامل التعيين
challengeType: 1
forumTopicId: 301171
dashedName: initializing-variables-with-the-assignment-operator
---

# --description--

من الشائع <dfn>تهيئة</dfn> متغير بقيمة أولية في نفس السطر الذي يُعلن فيه.

```js
var myVar = 0;
```

ينشئ متغيرًا جديدًا اسمه `myVar` ويُعيّن له قيمة أولية تساوي `0`.

# --instructions--

أعلن عن متغير `a` باستخدام `var` وتهيئته بقيمة `9`.

# --hints--

يجب أن تهيئ `a` بقيمة `9`.

```js
assert(/var\s+a\s*=\s*9(\s*;?\s*)$/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
var a = 9;
```
