---
id: 56533eb9ac21ba0edf2244bc
title: قائمة التسوق
challengeType: 1
forumTopicId: 18280
dashedName: shopping-list
---

# --description--

أنشئ قائمة تسوق في المتغير `myList`. يجب أن تكون القائمة مصفوفة متعددة الأبعاد تحتوي على عدة مصفوفات فرعية.

يجب أن يحتوي العنصر الأول في كل مصفوفة فرعية على سلسلة نصية باسم العنصر. يجب أن يكون العنصر الثاني رقمًا يمثل الكمية، أي.

```js
["Chocolate Bar", 15]
```

يجب أن تحتوي القائمة على 5 مصفوفات فرعية على الأقل.

# --hints--

يجب أن يكون `myList` مصفوفة.

```js
assert(Array.isArray(myList));
```

يجب أن تكون العناصر الأولى في كل مصفوفة فرعية من نوع سلسلة نصية.

```js
assert(myList.every(elem => elem && typeof elem[0] === 'string'));
```

يجب أن تكون العناصر الثانية في كل مصفوفة فرعية أرقامًا.

```js
assert(myList.every(elem => elem && typeof elem[1] === 'number'));
```

يجب أن تحتوي قائمتك على 5 عناصر على الأقل.

```js
assert(myList.length > 4);
```

# --seed--

## --seed-contents--

```js
const myList = [];
```

# --solutions--

```js
const myList = [
  ["Candy", 10],
  ["Potatoes", 12],
  ["Eggs", 12],
  ["Catfood", 1],
  ["Toads", 9]
];
```
