---
id: 56533eb9ac21ba0edf2244d3
title: المقارنة باستخدام معامل عدم المساواة الصارم
challengeType: 1
forumTopicId: 16791
dashedName: comparison-with-the-strict-inequality-operator
---

# --description--

معامل عدم المساواة الصارم (`!==`) هو العكس المنطقي لمعامل المساواة الصارمة. يعني "غير متساوٍ تمامًا" ويُرجع `false` حيث يُرجع معامل المساواة الصارمة `true` والعكس صحيح. لا يحول معامل عدم المساواة الصارم أنواع البيانات.

**أمثلة**

```js
3 !==  3  // false
3 !== '3' // true
4 !==  3  // true
```

# --instructions--

أضف معامل عدم المساواة الصارم إلى جملة `if` لكي تُرجع الدالة السلسلة النصية `Not Equal` عندما لا يكون `val` مساويًا تمامًا لـ `17`

# --hints--

`testStrictNotEqual(17)` يجب أن تُرجع السلسلة النصية `Equal`

```js
assert(testStrictNotEqual(17) === 'Equal');
```

`testStrictNotEqual("17")` يجب أن تُرجع السلسلة النصية `Not Equal`

```js
assert(testStrictNotEqual('17') === 'Not Equal');
```

`testStrictNotEqual(12)` يجب أن تُرجع السلسلة النصية `Not Equal`

```js
assert(testStrictNotEqual(12) === 'Not Equal');
```

`testStrictNotEqual("bob")` يجب أن تُرجع السلسلة النصية `Not Equal`

```js
assert(testStrictNotEqual('bob') === 'Not Equal');
```

يجب أن تستخدم معامل `!==`

```js
assert(__helpers.removeJSComments(code).match(/(val\s*!==\s*\d+)|(\d+\s*!==\s*val)/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function testStrictNotEqual(val) {
  if (val) { // Change this line
    return "Not Equal";
  }
  return "Equal";
}

testStrictNotEqual(10);
```

# --solutions--

```js
function testStrictNotEqual(val) {
  if (val !== 17) {
    return "Not Equal";
  }
  return "Equal";
}
```
