---
id: 56533eb9ac21ba0edf2244b2
title: تعيين مركب مع القسمة المعززة
challengeType: 1
forumTopicId: 16659
dashedName: compound-assignment-with-augmented-division
---

# --description--

المشغل `/=` يقسم متغيرًا على رقم آخر.

```js
myVar = myVar / 5;
```

سيقسم `myVar` على `5`. يمكن إعادة كتابته كالتالي:

```js
myVar /= 5;
```

# --instructions--

حوّل التعيينات لـ `a` و `b` و `c` لاستخدام المشغل `/=`.

# --hints--

يجب أن يساوي `a` القيمة `4`.

```js
assert(a === 4);
```

يجب أن يساوي `b` القيمة `27`.

```js
assert(b === 27);
```

يجب أن يساوي `c` القيمة `3`.

```js
assert(c === 3);
```

يجب استخدام المشغل `/=` لكل متغير.

```js
assert(__helpers.removeJSComments(code).match(/\/=/g).length === 3);
```

لا تعدل الشفرة فوق التعليق المحدد.

```js
assert(
  /let a = 48;/.test(__helpers.removeJSComments(code)) &&
    /let b = 108;/.test(__helpers.removeJSComments(code)) &&
    /let c = 33;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 48;
let b = 108;
let c = 33;

// Only change code below this line
a = a / 12;
b = b / 4;
c = c / 11;
```

# --solutions--

```js
let a = 48;
let b = 108;
let c = 33;

a /= 12;
b /= 4;
c /= 11;
```
