---
id: 56533eb9ac21ba0edf2244b6
title: متتاليات التخطي في السلاسل النصية
challengeType: 1
forumTopicId: 17567
dashedName: escape-sequences-in-strings
---

# --description--

ليست علامات الاقتباس هي الوحيدة التي يمكن <dfn>تخطيها</dfn> داخل سلسلة نصية. تتيح لك متتاليات التخطي استخدام أحرف قد لا تتمكن من استخدامها في السلسلة النصية بخلاف ذلك.

<table><thead><tr><th>الشفرة</th><th>الناتج</th></tr></thead><tbody><tr><td><code>\'</code></td><td>علامة اقتباس مفردة</td></tr><tr><td><code>\"</code></td><td>علامة اقتباس مزدوجة</td></tr><tr><td><code>\\</code></td><td>شرطة مائلة عكسية</td></tr><tr><td><code>\n</code></td><td>سطر جديد</td></tr><tr><td><code>\t</code></td><td>تاب</td></tr><tr><td><code>\r</code></td><td>عودة المؤشر إلى بداية السطر</td></tr><tr><td><code>\b</code></td><td>مسافة خلفية</td></tr><tr><td><code>\f</code></td><td>تغذية النموذج</td></tr></tbody></table>

*لاحظ أنه يجب تخطي الشرطة المائلة العكسية نفسها لكي تظهر كشرطة مائلة عكسية.*

# --instructions--

عيّن الأسطر الثلاثة التالية من النص في المتغير الواحد `myStr` باستخدام متتاليات التخطي.

<pre>
FirstLine
    \SecondLine
ThirdLine
</pre>

ستحتاج إلى استخدام متتاليات التخطي لإدخال الأحرف الخاصة بشكل صحيح. كما يجب أن تتبع التباعد كما هو موضح أعلاه، دون وجود مسافات بين متتاليات التخطي أو الكلمات.

**ملاحظة:** التراجع في السطر `SecondLine` يتم باستخدام حرف التاب في متتالية التخطي، وليس باستخدام المسافات.

# --hints--

يجب ألا يحتوي `myStr` على أي مسافات

```js
assert(!/ /.test(myStr));
```

يجب أن يحتوي `myStr` على السلاسل النصية `FirstLine` و`SecondLine` و`ThirdLine` (تذكر حساسية الحروف)

```js
assert(
  /FirstLine/.test(myStr) && /SecondLine/.test(myStr) && /ThirdLine/.test(myStr)
);
```

يجب أن يتبع `FirstLine` حرف السطر الجديد `\n`

```js
assert(/FirstLine\n/.test(myStr));
```

يجب أن يحتوي `myStr` على حرف تاب `\t` يتبع حرف السطر الجديد

```js
assert(/\n\t/.test(myStr));
```

يجب أن يسبق `SecondLine` حرف الشرطة المائلة العكسية `\`

```js
assert(/\\SecondLine/.test(myStr));
```

يجب أن يكون هناك حرف سطر جديد بين `SecondLine` و`ThirdLine`

```js
assert(/SecondLine\nThirdLine/.test(myStr));
```

يجب أن يحتوي `myStr` فقط على الأحرف الموضحة في التعليمات

```js
assert(myStr === 'FirstLine\n\t\\SecondLine\nThirdLine');
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "FirstLine\n\t\\SecondLine\nThirdLine";
```
