---
id: 56533eb9ac21ba0edf2244d9
title: المقارنات باستخدام عامل أو المنطقي
challengeType: 1
forumTopicId: 16800
dashedName: comparisons-with-the-logical-or-operator
---

# --description--

عامل <dfn>أو المنطقي</dfn> (`||`) يُرجع `true` إذا كان أي من <dfn>العوامل</dfn> يساوي `true`. وإلا، يُرجع `false`.

عامل <dfn>أو المنطقي</dfn> يتكون من رمزين للأنابيب: (`||`). عادةً ما تجده بين مفتاحي Backspace و Enter.

النمط أدناه يجب أن يكون مألوفًا من نقاط التوقف السابقة.

```js
if (num > 10) {
  return "No";
}
if (num < 5) {
  return "No";
}
return "Yes";
```

هذه الشفرة سترجع `Yes` إذا كان `num` بين `5` و `10` (شاملًا `5` و `10`). يمكن كتابة نفس المنطق باستخدام عامل <dfn>أو المنطقي</dfn>.

```js
if (num > 10 || num < 5) {
  return "No";
}
return "Yes";
```

# --instructions--

ادمج جملتي `if` في جملة واحدة تُرجع السلسلة النصية `Outside` إذا لم يكن `val` بين `10` و `20`، شاملاً القيمتين. وإلا، أرجع السلسلة النصية `Inside`.

# --hints--

يجب أن تستخدم عامل `||` مرة واحدة

```js
assert(__helpers.removeJSComments(code).match(/\|\|/g).length === 1);
```

يجب أن يكون لديك جملة `if` واحدة فقط

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`testLogicalOr(0)` يجب أن تُرجع السلسلة النصية `Outside`

```js
assert(testLogicalOr(0) === 'Outside');
```

`testLogicalOr(9)` يجب أن تُرجع السلسلة النصية `Outside`

```js
assert(testLogicalOr(9) === 'Outside');
```

`testLogicalOr(10)` يجب أن تُرجع السلسلة النصية `Inside`

```js
assert(testLogicalOr(10) === 'Inside');
```

`testLogicalOr(15)` يجب أن تُرجع السلسلة النصية `Inside`

```js
assert(testLogicalOr(15) === 'Inside');
```

`testLogicalOr(19)` يجب أن تُرجع السلسلة النصية `Inside`

```js
assert(testLogicalOr(19) === 'Inside');
```

`testLogicalOr(20)` يجب أن تُرجع السلسلة النصية `Inside`

```js
assert(testLogicalOr(20) === 'Inside');
```

`testLogicalOr(21)` يجب أن تُرجع السلسلة النصية `Outside`

```js
assert(testLogicalOr(21) === 'Outside');
```

`testLogicalOr(25)` يجب أن تُرجع السلسلة النصية `Outside`

```js
assert(testLogicalOr(25) === 'Outside');
```

# --seed--

## --seed-contents--

```js
function testLogicalOr(val) {
  // Only change code below this line

  if (val) {
    return "Outside";
  }

  if (val) {
    return "Outside";
  }

  // Only change code above this line
  return "Inside";
}

testLogicalOr(15);
```

# --solutions--

```js
function testLogicalOr(val) {
  if (val < 10 || val > 20) {
    return "Outside";
  }
  return "Inside";
}
```
