---
id: 56533eb9ac21ba0edf2244a8
title: تخزين القيم باستخدام معامل التعيين
challengeType: 1
forumTopicId: 18310
dashedName: storing-values-with-the-assignment-operator
---

# --description--

في JavaScript، يمكنك تخزين قيمة في متغير باستخدام <dfn>معامل التعيين</dfn> (`=`).

```js
myVariable = 5;
```

هذا يُعيّن القيمة `5` من نوع `Number` إلى `myVariable`.

إذا وُجدت أي عمليات حسابية على يمين معامل `=`، تُنفذ أولًا قبل تعيين القيمة إلى المتغير على يسار المعامل.

```js
var myVar;
myVar = 5;
```

أولًا، ينشئ هذا الشفرة متغيرًا باسم `myVar`. ثم تُعيّن القيمة `5` إلى `myVar`. الآن، إذا ظهر `myVar` مرة أخرى في الشفرة، سيعامل البرنامج قيمته كما لو كانت `5`.

# --instructions--

عيّن القيمة `7` إلى المتغير `a`.

# --hints--

لا يجب تغيير الشفرة فوق التعليق المحدد.

```js
assert(/var a;/.test(__helpers.removeJSComments(code)));
```

يجب أن تكون قيمة `a` هي 7.

```js
assert(typeof a === 'number' && a === 7);
```

# --seed--

## --seed-contents--

```js
// Setup
var a;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
```
