---
id: 56533eb9ac21ba0edf2244ae
title: إيجاد الباقي في JavaScript
challengeType: 1
forumTopicId: 18184
dashedName: finding-a-remainder-in-javascript
---

# --description--

عامل <dfn>الباقي</dfn> `%` يعطي الباقي من قسمة رقمين.

**مثال**

<pre>
5 % 2 = 1
5 / 2 = 2 الباقي 1
2 * 2 = 4
5 - 4 = 1
</pre>

**الاستخدام**  
في الرياضيات، يمكن التحقق مما إذا كان الرقم زوجيًا أو فرديًا بفحص الباقي من قسمة الرقم على `2`. الأرقام الزوجية يكون لها باقي `0`، بينما الأرقام الفردية يكون لها باقي `1`.

<pre>
17 % 2 = 1
48 % 2 = 0
</pre>

**ملاحظة:** يُشار أحيانًا إلى عامل <dfn>الباقي</dfn> بشكل خاطئ على أنه عامل معامل باقي القسمة. هو مشابه جدًا لمعامل باقي القسمة، لكنه لا يعمل بشكل صحيح مع الأعداد السالبة.

# --instructions--

اجعل قيمة `remainder` تساوي الباقي من قسمة `11` على `3` باستخدام عامل <dfn>الباقي</dfn> (`%`).

# --hints--

يجب تهيئة المتغير `remainder`

```js
assert(/(const|let|var)\s+?remainder/.test(__helpers.removeJSComments(code)));
```

يجب أن تكون قيمة `remainder` هي `2`

```js
assert(remainder === 2);
```

يجب استخدام عامل `%`

```js
assert(/\s+?remainder\s*?=\s*?.*%.*;?/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const remainder = 0;
```

# --solutions--

```js
const remainder = 11 % 3;
```
