---
id: 56533eb9ac21ba0edf2244c9
title: الوصول إلى خصائص الكائن باستخدام المتغيرات
challengeType: 1
forumTopicId: 16165
dashedName: accessing-object-properties-with-variables
---

# --description--

استخدام آخر لترميز الأقواس على الكائنات هو الوصول إلى خاصية مخزنة كقيمة لمتغير. هذا يكون مفيدًا جدًا عند التكرار عبر خصائص الكائن أو عند الوصول إلى جدول بحث.

إليك مثالًا على استخدام متغير للوصول إلى خاصية:

```js
const dogs = {
  Fido: "Mutt",
  Hunter: "Doberman",
  Snoopie: "Beagle"
};

const myDog = "Hunter";
const myBreed = dogs[myDog];
console.log(myBreed);
```

سوف تُعرض السلسلة النصية `Doberman` في وحدة التحكم.

لاحظ أننا *لا* نستخدم علامات اقتباس حول اسم المتغير عند استخدامه للوصول إلى الخاصية لأننا نستخدم *قيمة* المتغير، وليس *الاسم*.

# --instructions--

عيّن المتغير `playerNumber` إلى `16`. ثم استخدم المتغير للبحث عن اسم اللاعب وعيّنه إلى `player`.

# --hints--

يجب أن يكون `playerNumber` رقمًا

```js
assert(typeof playerNumber === 'number');
```

يجب أن يكون المتغير `player` سلسلة نصية

```js
assert(typeof player === 'string');
```

يجب أن تكون قيمة `player` هي السلسلة النصية `Montana`

```js
assert(player === 'Montana');
```

يجب أن تستخدم ترميز الأقواس للوصول إلى `testObj`

```js
assert(/testObj\s*?\[.*?\]/.test(__helpers.removeJSComments(code)));
```

لا يجب أن تعيّن القيمة `Montana` إلى المتغير `player` مباشرةً.

```js
assert(!__helpers.removeJSComments(code).match(/player\s*=\s*"|\'\s*Montana\s*"|\'\s*;/gi));
```

يجب أن تستخدم المتغير `playerNumber` في ترميز الأقواس الخاص بك

```js
assert(/testObj\s*?\[\s*playerNumber\s*\]/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  12: "Namath",
  16: "Montana",
  19: "Unitas"
};

// Only change code below this line
const playerNumber = 42;  // Change this line
const player = testObj;   // Change this line
```

# --solutions--

```js
const testObj = {
  12: "Namath",
  16: "Montana",
  19: "Unitas"
};
const playerNumber = 16;
const player = testObj[playerNumber];
```
