---
id: 56533eb9ac21ba0edf2244dd
title: الاختيار من بين العديد من الخيارات باستخدام جمل التبديل
challengeType: 1
forumTopicId: 18277
dashedName: selecting-from-many-options-with-switch-statements
---

# --description--

إذا كنت تحتاج إلى مطابقة قيمة واحدة مع العديد من الخيارات، يمكنك استخدام جملة <dfn>switch</dfn>. تقارن جملة `switch` القيمة مع جمل <dfn>case</dfn> التي تحدد القيم المحتملة المختلفة. يمكن تنفيذ أي جمل JavaScript صحيحة داخل كتلة <dfn>case</dfn> وستُنفذ بدءًا من أول قيمة `case` مطابقة حتى يتم الوصول إلى `break`.

إليك مثالًا على جملة `switch`:

```js
switch (fruit) {
  case "apple":
    console.log("The fruit is an apple");
    break;
  case "orange":
    console.log("The fruit is an orange");
    break;
}
```

تُختبر قيم `case` بالمساواة الصارمة (`===`). يخبر `break` جافاسكريبت بالتوقف عن تنفيذ الجمل. إذا حُذفت `break`، ستُنفذ الجملة التالية.

# --instructions--

اكتب جملة switch تختبر `val` وتعيّن `answer` للحالات التالية:  
`1` - `alpha`  
`2` - `beta`  
`3` - `gamma`  
`4` - `delta`

# --hints--

يجب أن تكون قيمة `caseInSwitch(1)` هي السلسلة النصية `alpha`

```js
assert(caseInSwitch(1) === 'alpha');
```

يجب أن تكون قيمة `caseInSwitch(2)` هي السلسلة النصية `beta`

```js
assert(caseInSwitch(2) === 'beta');
```

يجب أن تكون قيمة `caseInSwitch(3)` هي السلسلة النصية `gamma`

```js
assert(caseInSwitch(3) === 'gamma');
```

يجب أن تكون قيمة `caseInSwitch(4)` هي السلسلة النصية `delta`

```js
assert(caseInSwitch(4) === 'delta');
```

يجب ألا تستخدم أي جمل `if` أو `else`

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

يجب أن تحتوي على 3 جمل `break` على الأقل

```js
assert(__helpers.removeJSComments(code).match(/break/g).length > 2);
```

# --seed--

## --seed-contents--

```js
function caseInSwitch(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

caseInSwitch(1);
```

# --solutions--

```js
function caseInSwitch(val) {
  let answer = "";

  switch (val) {
    case 1:
      answer = "alpha";
      break;
    case 2:
      answer = "beta";
      break;
    case 3:
      answer = "gamma";
      break;
    case 4:
      answer = "delta";
  }
  return answer;
}
```
