---
id: 56533eb9ac21ba0edf2244d1
title: المقارنة باستخدام معامل المساواة الصارمة
challengeType: 1
forumTopicId: 16790
dashedName: comparison-with-the-strict-equality-operator
---

# --description--

معامل المساواة الصارمة (`===`) هو النظير لمعامل المساواة (`==`). ومع ذلك، على عكس معامل المساواة الذي يحاول تحويل القيمتين المقارنتين إلى نوع مشترك، لا يُجري معامل المساواة الصارمة تحويل نوع.

إذا كانت القيم المقارنة من أنواع مختلفة، تُعتبر غير متساوية، وسيُرجع معامل المساواة الصارمة القيمة false.

**أمثلة**

```js
3 ===  3  // true
3 === '3' // false
```

في المثال الثاني، القيمة `3` من نوع `Number` والقيمة `'3'` من نوع `String`.

# --instructions--

استخدم معامل المساواة الصارمة في جملة `if` لكي تُرجع الدالة السلسلة النصية `Equal` عندما تكون `val` مساوية تمامًا لـ `7`.

# --hints--

يجب أن تُرجع الدالة `testStrict(10)` السلسلة النصية `Not Equal`

```js
assert(testStrict(10) === 'Not Equal');
```

يجب أن تُرجع الدالة `testStrict(7)` السلسلة النصية `Equal`

```js
assert(testStrict(7) === 'Equal');
```

يجب أن تُرجع الدالة `testStrict("7")` السلسلة النصية `Not Equal`

```js
assert(testStrict('7') === 'Not Equal');
```

يجب أن تستخدم معامل `===`

```js
assert(__helpers.removeJSComments(code).match(/(val\s*===\s*\d+)|(\d+\s*===\s*val)/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function testStrict(val) {
  if (val) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

testStrict(10);
```

# --solutions--

```js
function testStrict(val) {
  if (val === 7) {
    return "Equal";
  }
  return "Not Equal";
}
```
