---
id: 56533eb9ac21ba0edf2244d7
title: المقارنة باستخدام معامل أصغر من أو يساوي
challengeType: 1
forumTopicId: 16788
dashedName: comparison-with-the-less-than-or-equal-to-operator
---

# --description--

يقارن معامل أصغر من أو يساوي (`<=`) بين قيم عددين. إذا كان العدد على اليسار أصغر من أو يساوي العدد على اليمين، يُرجع `true`. إذا كان العدد على اليسار أكبر من العدد على اليمين، يُرجع `false`. مثل معامل المساواة، يحول معامل أصغر من أو يساوي أنواع البيانات.

**أمثلة**

```js
4   <= 5 // true
'7' <= 7 // true
5   <= 5 // true
3   <= 2 // false
'8' <= 4 // false
```

# --instructions--

أضف معامل أصغر من أو يساوي إلى الأسطر المحددة بحيث تكون جمل الإرجاع منطقية.

# --hints--

`testLessOrEqual(0)` يجب أن تُرجع السلسلة النصية `Smaller Than or Equal to 12`

```js
assert(testLessOrEqual(0) === 'Smaller Than or Equal to 12');
```

`testLessOrEqual(11)` يجب أن تُرجع السلسلة النصية `Smaller Than or Equal to 12`

```js
assert(testLessOrEqual(11) === 'Smaller Than or Equal to 12');
```

`testLessOrEqual(12)` يجب أن تُرجع السلسلة النصية `Smaller Than or Equal to 12`

```js
assert(testLessOrEqual(12) === 'Smaller Than or Equal to 12');
```

`testLessOrEqual(23)` يجب أن تُرجع السلسلة النصية `Smaller Than or Equal to 24`

```js
assert(testLessOrEqual(23) === 'Smaller Than or Equal to 24');
```

`testLessOrEqual(24)` يجب أن تُرجع السلسلة النصية `Smaller Than or Equal to 24`

```js
assert(testLessOrEqual(24) === 'Smaller Than or Equal to 24');
```

`testLessOrEqual(25)` يجب أن تُرجع السلسلة النصية `More Than 24`

```js
assert(testLessOrEqual(25) === 'More Than 24');
```

`testLessOrEqual(55)` يجب أن تُرجع السلسلة النصية `More Than 24`

```js
assert(testLessOrEqual(55) === 'More Than 24');
```

يجب أن تستخدم معامل `<=` مرتين على الأقل

```js
assert(__helpers.removeJSComments(code).match(/val\s*<=\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testLessOrEqual(val) {
  if (val) {  // Change this line
    return "Smaller Than or Equal to 12";
  }

  if (val) {  // Change this line
    return "Smaller Than or Equal to 24";
  }

  return "More Than 24";
}

testLessOrEqual(10);
```

# --solutions--

```js
function testLessOrEqual(val) {
  if (val <= 12) {  // Change this line
    return "Smaller Than or Equal to 12";
  }

  if (val <= 24) {  // Change this line
    return "Smaller Than or Equal to 24";
  }

  return "More Than 24";
}
```
