---
id: 56533eb9ac21ba0edf2244cf
title: مجموعة التسجيلات
challengeType: 1
forumTopicId: 18261
dashedName: record-collection
---

# --description--

أنت تنشئ دالة تساعد في صيانة مجموعة ألبومات موسيقية. تُنظم المجموعة ككائن يحتوي على عدة ألبومات هي أيضًا كائنات. يُمثل كل ألبوم في المجموعة بمعرف `id` فريد كاسم خاصية. داخل كل كائن ألبوم، توجد خصائص مختلفة تصف معلومات عن الألبوم. ليست كل الألبومات تحتوي على معلومات كاملة.

تأخذ الدالة `updateRecords` أربعة معلمات ممثلة بالمعلمات التالية للدالة:

-   `records` - كائن يحتوي على عدة ألبومات فردية
-   `id` - رقم يمثل ألبومًا معينًا في كائن `records`
-   `prop` - سلسلة نصية تمثل اسم الخاصية في الألبوم التي سيتم تحديثها
-   `value` - سلسلة نصية تحتوي على المعلومات المستخدمة لتحديث خاصية الألبوم

أكمل الدالة باستخدام القواعد التالية لتعديل الكائن الممرر إلى الدالة.

-   يجب أن تُعيد الدالة دائمًا كائن `records` بالكامل.
-   إذا كانت `value` سلسلة نصية فارغة، فاحذف الخاصية `prop` المحددة من الألبوم.
-   إذا لم تكن `prop` هي `tracks` ولم تكن `value` سلسلة نصية فارغة، فعيّن `value` إلى `prop` في ذلك الألبوم.
-   إذا كانت `prop` هي `tracks` ولم تكن `value` سلسلة نصية فارغة، لكن الألبوم لا يحتوي على الخاصية `tracks`، فأنشئ مصفوفة فارغة وأضف `value` إليها.
-   إذا كانت prop هي `tracks` ولم تكن `value` سلسلة نصية فارغة، فأضف `value` إلى نهاية مصفوفة `tracks` الموجودة في الألبوم.

**ملاحظة:** تُستخدم نسخة من كائن `recordCollection` للاختبارات. لا يجب تعديل كائن `recordCollection` مباشرة.

# --before-each--

```js
const _recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};
```

# --hints--

بعد `updateRecords(recordCollection, 5439, "artist", "ABBA")`، يجب أن تكون الخاصية `artist` هي السلسلة النصية `ABBA`

```js
assert(
  updateRecords(_recordCollection, 5439, 'artist', 'ABBA')[5439]['artist'] ===
    'ABBA'
);
```

بعد `updateRecords(recordCollection, 5439, "tracks", "Take a Chance on Me")`، يجب أن تحتوي الخاصية `tracks` على السلسلة النصية `Take a Chance on Me` كالعنصر الأخير والوحيد.

```js
assert(
  updateRecords(_recordCollection, 5439, 'tracks', 'Take a Chance on Me') &&
  _recordCollection[5439]['tracks'].length === 1 &&
  _recordCollection[5439]['tracks'].pop() === 'Take a Chance on Me'
);
```

بعد `updateRecords(recordCollection, 2548, "artist", "")`، يجب ألا تُعيّن الخاصية `artist`

```js
updateRecords(_recordCollection, 2548, 'artist', '');
assert(!_recordCollection[2548].hasOwnProperty('artist'));
```

بعد `updateRecords(recordCollection, 1245, "tracks", "Addicted to Love")`، يجب أن تحتوي الخاصية `tracks` على السلسلة النصية `Addicted to Love` كالعنصر الأخير.

```js
assert(
  updateRecords(_recordCollection, 1245, 'tracks', 'Addicted to Love')[1245][
    'tracks'
  ].pop() === 'Addicted to Love'
);
```

بعد `updateRecords(recordCollection, 2468, "tracks", "Free")`، يجب أن تحتوي الخاصية `tracks` على السلسلة النصية `1999` كالعنصر الأول.

```js
assert(
  updateRecords(_recordCollection, 2468, 'tracks', 'Free')[2468][
    'tracks'
  ][0] === '1999'
);
```

بعد `updateRecords(recordCollection, 2548, "tracks", "")`، يجب ألا تُعيّن الخاصية `tracks`

```js
updateRecords(_recordCollection, 2548, 'tracks', '');
assert(!_recordCollection[2548].hasOwnProperty('tracks'));
```

بعد `updateRecords(recordCollection, 1245, "albumTitle", "Riptide")`، يجب أن تكون الخاصية `albumTitle` هي السلسلة النصية `Riptide`

```js
assert(
  updateRecords(_recordCollection, 1245, 'albumTitle', 'Riptide')[1245][
    'albumTitle'
  ] === 'Riptide'
);
```

# --seed--

## --seed-contents--

```js
// Setup
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};

// Only change code below this line
function updateRecords(records, id, prop, value) {
  return records;
}

updateRecords(recordCollection, 5439, 'artist', 'ABBA');
```

# --solutions--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};

// Only change code below this line
function updateRecords(records, id, prop, value) {
  if (value === '') delete records[id][prop];
  else if (prop === 'tracks') {
    records[id][prop] = records[id][prop] || [];
    records[id][prop].push(value);
  } else {
    records[id][prop] = value;
  }

  return records;
}
```
