---
id: 56533eb9ac21ba0edf2244aa
title: فهم المتغيرات غير المهيأة
challengeType: 1
forumTopicId: 18335
dashedName: understanding-uninitialized-variables
---

# --description--

عند إعلان متغيرات JavaScript، يكون لها قيمة ابتدائية `undefined`. إذا أجريت عملية حسابية على متغير `undefined` ستكون النتيجة `NaN` والتي تعني <dfn>"ليست رقمًا"</dfn>. إذا دمجت سلسلة نصية مع متغير `undefined` ستحصل على <dfn>سلسلة نصية</dfn> تحتوي على `undefined`.

# --instructions--

هيئ المتغيرات الثلاثة `a` و `b` و `c` بالقيم `5` و `10` و `"I am a"` على التوالي بحيث لا تكون قيمتها `undefined`.

# --hints--

يجب أن يكون `a` معرفًا وله القيمة النهائية `6`.

```js
assert(typeof a === 'number' && a === 6);
```

يجب أن يكون `b` معرفًا وله القيمة النهائية `15`.

```js
assert(typeof b === 'number' && b === 15);
```

يجب ألا يحتوي `c` على `undefined` ويجب أن تكون قيمته النهائية سلسلة نصية `I am a String!`

```js
assert(!/undefined/.test(c) && c === 'I am a String!');
```

لا يجب تغيير الشفرة أسفل التعليق المحدد.

```js
assert(
  /a = a \+ 1;/.test(__helpers.removeJSComments(code)) &&
    /b = b \+ 5;/.test(__helpers.removeJSComments(code)) &&
    /c = c \+ " String!";/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
var a;
var b;
var c;
// Only change code above this line

a = a + 1;
b = b + 5;
c = c + " String!";
```

# --solutions--

```js
var a = 5;
var b = 10;
var c = "I am a";
a = a + 1;
b = b + 5;
c = c + " String!";
```
