---
id: 56533eb9ac21ba0edf2244c7
title: الوصول إلى خصائص الكائن باستخدام ترميز النقطة
challengeType: 1
forumTopicId: 16164
dashedName: accessing-object-properties-with-dot-notation
---

# --description--

هناك طريقتان للوصول إلى خصائص الكائن: ترميز النقطة (`.`) وترميز الأقواس (`[]`)، يشبهان المصفوفة.

ترميز النقطة تستخدمه عندما تعرف اسم الخاصية التي تريد الوصول إليها مسبقًا.

إليك مثالًا على استخدام ترميز النقطة (`.`) لقراءة خاصية من كائن:

```js
const myObj = {
  prop1: "val1",
  prop2: "val2"
};

const prop1val = myObj.prop1;
const prop2val = myObj.prop2;
```

سيكون لـ `prop1val` قيمة السلسلة النصية `val1`، ولـ `prop2val` قيمة السلسلة النصية `val2`.

# --instructions--

اقرأ قيم الخصائص من `testObj` باستخدام ترميز النقطة. عرّف المتغير `hatValue` مساويًا لخاصية `hat` في الكائن، والمتغير `shirtValue` مساويًا لخاصية `shirt` في الكائن.

# --hints--

يجب أن يكون `hatValue` سلسلة نصية

```js
assert(typeof hatValue === 'string');
```

يجب أن تكون قيمة `hatValue` هي السلسلة النصية `ballcap`

```js
assert(hatValue === 'ballcap');
```

يجب أن يكون `shirtValue` سلسلة نصية

```js
assert(typeof shirtValue === 'string');
```

يجب أن تكون قيمة `shirtValue` هي السلسلة النصية `jersey`

```js
assert(shirtValue === 'jersey');
```

يجب أن تستخدم ترميز النقطة مرتين

```js
assert(__helpers.removeJSComments(code).match(/testObj\.\w+/g).length > 1);
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  hat: "ballcap",
  shirt: "jersey",
  shoes: "cleats"
};

// Only change code below this line
const hatValue = testObj;      // Change this line
const shirtValue = testObj;    // Change this line
```

# --solutions--

```js
const testObj = {
  hat: "ballcap",
  shirt: "jersey",
  shoes: "cleats"
};

const hatValue = testObj.hat;
const shirtValue = testObj.shirt;
```
