---
id: 56533eb9ac21ba0edf2244be
title: النطاق العام والدوال
challengeType: 1
forumTopicId: 18193
dashedName: global-scope-and-functions
---

# --description--

في JavaScript، يشير <dfn>النطاق</dfn> إلى مستوى الظهور للمتغيرات. المتغيرات التي تُعرّف خارج كتلة دالة لها نطاق <dfn>عام</dfn>. هذا يعني أنها تُرى في كل مكان في شفرة JavaScript الخاصة بك.

المتغيرات التي تُعلن بدون الكلمات المفتاحية `let` أو `const` تُنشأ تلقائيًا في النطاق `global`. هذا قد يسبب نتائج غير مقصودة في أماكن أخرى من شفرتك أو عند تشغيل دالة مرة أخرى. يجب دائمًا إعلان متغيراتك باستخدام `let` أو `const`.

# --instructions--

باستخدام `let` أو `const`، أعلن متغيرًا عامًا باسم `myGlobal` خارج أي دالة. ابدأه بقيمة `10`.

داخل الدالة `fun1`، عيّن القيمة `5` للمتغير `oopsGlobal` ***بدون*** استخدام الكلمات المفتاحية `var` أو `let` أو `const`.

# --before-each--

```js
var oopsGlobal;
```

# --hints--

يجب تعريف `myGlobal`

```js
assert(typeof myGlobal != 'undefined');
```

يجب أن تكون قيمة `myGlobal` هي `10`

```js
assert(myGlobal === 10);
```

يجب إعلان `myGlobal` باستخدام الكلمات المفتاحية `let` أو `const`

```js
assert(/(let|const)\s+myGlobal/.test(__helpers.removeJSComments(code)));
```

يجب أن يكون `oopsGlobal` متغيرًا عامًا وله قيمة `5`

```js
fun1();
assert(typeof oopsGlobal != 'undefined');
```

# --seed--

## --seed-contents--

```js
// Declare the myGlobal variable below this line


function fun1() {
  // Assign 5 to oopsGlobal here

}

// Only change code above this line

function fun2() {
  let output = "";
  if (typeof myGlobal != "undefined") {
    output += "myGlobal: " + myGlobal;
  }
  if (typeof oopsGlobal != "undefined") {
    output += " oopsGlobal: " + oopsGlobal;
  }
  console.log(output);
}
```

# --solutions--

```js
const myGlobal = 10;

function fun1() {
  oopsGlobal = 5;
}

function fun2() {
  let output = "";
  if(typeof myGlobal != "undefined") {
    output += "myGlobal: " + myGlobal;
  }
  if(typeof oopsGlobal != "undefined") {
    output += " oopsGlobal: " + oopsGlobal;
  }
  console.log(output);
}
```
