---
id: 56533eb9ac21ba0edf2244b7
title: دمج السلاسل النصية باستخدام عامل الجمع
challengeType: 1
forumTopicId: 16802
dashedName: concatenating-strings-with-plus-operator
---

# --description--

في JavaScript، عندما يُستخدم العامل `+` مع قيمة من نوع `String`، يُسمى عامل <dfn>الدمج</dfn>. يمكنك بناء سلسلة نصية جديدة من سلاسل أخرى عن طريق <dfn>دمج</dfn> هذه السلاسل معًا.

**مثال**

```js
'My name is Alan,' + ' I concatenate.'
```

**ملاحظة:** انتبه للمسافات. الدمج لا يضيف مسافات بين السلاسل المدمجة، لذا عليك إضافتها بنفسك.

مثال:

```js
const ourStr = "I come first. " + "I come second.";
```

ستُعرض السلسلة النصية `I come first. I come second.` في وحدة التحكم.
# --instructions--

ابنِ `myStr` من السلاسل النصية `This is the start.` و `This is the end.` باستخدام العامل `+`. تأكد من تضمين مسافة بين السلسلتين.

# --hints--

يجب أن يحتوي `myStr` على حرف مسافة واحد بين السلسلتين.

```js
assert(/start\. This/.test(myStr));
```

يجب أن تكون قيمة `myStr` هي السلسلة النصية `This is the start. This is the end.`

```js
assert(myStr === 'This is the start. This is the end.');
```

يجب استخدام العامل `+` لبناء `myStr`.

```js
assert(__helpers.removeJSComments(code).match(/(["']).*\1\s*\+\s*(["']).*\2/g));
```

يجب إنشاء `myStr` باستخدام الكلمة المفتاحية `const`.

```js
assert(/const\s+myStr/.test(__helpers.removeJSComments(code)));
```

يجب تعيين النتيجة إلى المتغير `myStr`.

```js
assert(/myStr\s*=/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "This is the start. " + "This is the end.";
```
