---
id: 56533eb9ac21ba0edf2244d0
title: المقارنة باستخدام معامل المساواة
challengeType: 1
forumTopicId: 16784
dashedName: comparison-with-the-equality-operator
---

# --description--

هناك العديد من <dfn>معاملات المقارنة</dfn> في JavaScript. جميع هذه المعاملات تُرجع قيمة منطقية `true` أو `false`.

المعامل الأساسي هو معامل المساواة `==`. يقارن معامل المساواة بين قيمتين ويُرجع `true` إذا كانتا متكافئتين أو `false` إذا لم تكونا كذلك. لاحظ أن المساواة تختلف عن التعيين (`=`)، الذي يُعيّن القيمة على يمين المعامل إلى متغير على اليسار.

```js
function equalityTest(myVal) {
  if (myVal == 10) {
    return "Equal";
  }
  return "Not Equal";
}
```

إذا كانت قيمة `myVal` تساوي `10`، يُرجع معامل المساواة `true`، لذا تُنفذ الشفرة داخل الأقواس المتعرجة وتُرجع الدالة `Equal`. وإلا، تُرجع الدالة `Not Equal`. لكي يقارن JavaScript بين نوعين مختلفين من <dfn>البيانات</dfn> (مثل `numbers` و `strings`)، يجب أن يُحوّل أحد النوعين إلى الآخر. يُعرف هذا بـ إجبار النوع. بعد ذلك، يمكنه مقارنة القيم كما يلي:

```js
1   ==  1  // true
1   ==  2  // false
1   == '1' // true
"3" ==  3  // true
```

# --instructions--

أضف معامل المساواة إلى السطر المشار إليه بحيث تُرجع الدالة السلسلة النصية `Equal` عندما تكون `val` مكافئة لـ `12`.

# --hints--

`testEqual(10)` يجب أن تُرجع السلسلة النصية `Not Equal`

```js
assert(testEqual(10) === 'Not Equal');
```

`testEqual(12)` يجب أن تُرجع السلسلة النصية `Equal`

```js
assert(testEqual(12) === 'Equal');
```

`testEqual("12")` يجب أن تُرجع السلسلة النصية `Equal`

```js
assert(testEqual('12') === 'Equal');
```

يجب أن تستخدم المعامل `==`

```js
assert(__helpers.removeJSComments(code).match(/==/g) && !__helpers.removeJSComments(code).match(/===/g));
```

# --seed--

## --seed-contents--

```js
// Setup
function testEqual(val) {
  if (val) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

testEqual(10);
```

# --solutions--

```js
function testEqual(val) {
  if (val == 12) {
    return "Equal";
  }
  return "Not Equal";
}
```
