---
id: 56533eb9ac21ba0edf2244ab
title: فهم حساسية حالة الأحرف في المتغيرات
challengeType: 1
forumTopicId: 18334
dashedName: understanding-case-sensitivity-in-variables
---

# --description--

في JavaScript، جميع أسماء المتغيرات والدوال حساسة لحالة الأحرف. هذا يعني أن استخدام الحروف الكبيرة والصغيرة مهم.

`MYVAR` ليس هو نفسه `MyVar` ولا `myvar`. من الممكن وجود عدة متغيرات مميزة بنفس الاسم ولكن بحالات أحرف مختلفة. يُنصح بشدة، من أجل الوضوح، *بعدم* استخدام هذه الميزة في اللغة.

**أفضل الممارسات**

اكتب أسماء المتغيرات في JavaScript باستخدام <dfn>تنسيق سنام الجمل</dfn>. في <dfn>تنسيق سنام الجمل</dfn>، تحتوي أسماء المتغيرات متعددة الكلمات على الكلمة الأولى بأحرف صغيرة وأول حرف من كل كلمة لاحقة يكون كبيرًا.

**أمثلة:**

```js
var someVariable;
var anotherVariableName;
var thisVariableNameIsSoLong;
```

# --instructions--

عدّل الإعلانات والتعيينات الحالية بحيث تستخدم أسماؤها <dfn>تنسيق سنام الجمل</dfn>.

لا تنشئ أي متغيرات جديدة.

# --hints--

يجب تعريف `studlyCapVar` وأن تكون قيمته `10`.

```js
assert(typeof studlyCapVar !== 'undefined' && studlyCapVar === 10);
```

يجب تعريف `properCamelCase` وأن تكون قيمته سلسلة نصية `A String`.

```js
assert(
  typeof properCamelCase !== 'undefined' && properCamelCase === 'A String'
);
```

يجب تعريف `titleCaseOver` وأن تكون قيمته `9000`.

```js
assert(typeof titleCaseOver !== 'undefined' && titleCaseOver === 9000);
```

يجب أن يستخدم `studlyCapVar` تنسيق سنام الجمل في كل من قسم الإعلان وقسم التعيين.

```js
assert(__helpers.removeJSComments(code).match(/studlyCapVar/g).length === 2);
```

يجب أن يستخدم `properCamelCase` تنسيق سنام الجمل في كل من قسم الإعلان وقسم التعيين.

```js
assert(__helpers.removeJSComments(code).match(/properCamelCase/g).length === 2);
```

يجب أن يستخدم `titleCaseOver` تنسيق سنام الجمل في كل من قسم الإعلان وقسم التعيين.

```js
assert(__helpers.removeJSComments(code).match(/titleCaseOver/g).length === 2);
```

# --seed--

## --seed-contents--

```js
// Variable declarations
var StUdLyCapVaR;
var properCamelCase;
var TitleCaseOver;

// Variable assignments
STUDLYCAPVAR = 10;
PRoperCAmelCAse = "A String";
tITLEcASEoVER = 9000;
```

# --solutions--

```js
var studlyCapVar;
var properCamelCase;
var titleCaseOver;

studlyCapVar = 10;
properCamelCase = "A String";
titleCaseOver = 9000;
```
