---
id: 56533eb9ac21ba0edf2244c8
title: الوصول إلى خصائص الكائن باستخدام ترميز الأقواس
challengeType: 1
forumTopicId: 16163
dashedName: accessing-object-properties-with-bracket-notation
---

# --description--

الطريقة الثانية للوصول إلى خصائص الكائن هي ترميز الأقواس (`[]`). إذا كانت الخاصية التي تحاول الوصول إليها تحتوي على فراغ في اسمها، فستحتاج إلى استخدام ترميز الأقواس.

مع ذلك، يمكنك استخدام ترميز الأقواس على خصائص الكائن التي لا تحتوي على فراغات.

إليك مثالًا على استخدام ترميز الأقواس لقراءة خاصية من كائن:

```js
const myObj = {
  "Space Name": "Kirk",
  "More Space": "Spock",
  "NoSpace": "USS Enterprise"
};

myObj["Space Name"];
myObj['More Space'];
myObj["NoSpace"];
```

`myObj["Space Name"]` ستكون السلسلة النصية `Kirk`، و`myObj['More Space']` ستكون السلسلة النصية `Spock`، و`myObj["NoSpace"]` ستكون السلسلة النصية `USS Enterprise`.

لاحظ أن أسماء الخصائص التي تحتوي على فراغات يجب أن تكون بين علامات اقتباس (مفردة أو مزدوجة).

# --instructions--

اقرأ قيم الخاصيتين `an entree` و `the drink` من `testObj` باستخدام ترميز الأقواس وعيّنها إلى `entreeValue` و `drinkValue` على التوالي.

# --hints--

يجب أن تكون `entreeValue` سلسلة نصية

```js
assert(typeof entreeValue === 'string');
```

يجب أن تكون قيمة `entreeValue` هي السلسلة النصية `hamburger`

```js
assert(entreeValue === 'hamburger');
```

يجب أن تكون `drinkValue` سلسلة نصية

```js
assert(typeof drinkValue === 'string');
```

يجب أن تكون قيمة `drinkValue` هي السلسلة النصية `water`

```js
assert(drinkValue === 'water');
```

يجب أن تستخدم ترميز الأقواس مرتين

```js
assert(__helpers.removeJSComments(code).match(/testObj\s*?\[('|")[^'"]+\1\]/g).length > 1);
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  "an entree": "hamburger",
  "my side": "veggies",
  "the drink": "water"
};

// Only change code below this line
const entreeValue = testObj;   // Change this line
const drinkValue = testObj;    // Change this line
```

# --solutions--

```js
const testObj = {
  "an entree": "hamburger",
  "my side": "veggies",
  "the drink": "water"
};
const entreeValue = testObj["an entree"];
const drinkValue = testObj['the drink'];
```
