---
id: 56533eb9ac21ba0edf2244d6
title: المقارنة باستخدام معامل أصغر من
challengeType: 1
forumTopicId: 16789
dashedName: comparison-with-the-less-than-operator
---

# --description--

يقارن معامل أصغر من (`<`) بين قيم عددين. إذا كان العدد على اليسار أصغر من العدد على اليمين، يُرجع `true`. وإلا، يُرجع `false`. مثل معامل المساواة، يحول معامل أصغر من أنواع البيانات أثناء المقارنة.

**أمثلة**

```js
2   < 5 // true
'3' < 7 // true
5   < 5 // false
3   < 2 // false
'8' < 4 // false
```

# --instructions--

أضف معامل أصغر من إلى الأسطر المشار إليها بحيث تكون جمل الإرجاع منطقية.

# --hints--

`testLessThan(0)` يجب أن تُرجع السلسلة النصية `Under 25`

```js
assert(testLessThan(0) === 'Under 25');
```

`testLessThan(24)` يجب أن تُرجع السلسلة النصية `Under 25`

```js
assert(testLessThan(24) === 'Under 25');
```

`testLessThan(25)` يجب أن تُرجع السلسلة النصية `Under 55`

```js
assert(testLessThan(25) === 'Under 55');
```

`testLessThan(54)` يجب أن تُرجع السلسلة النصية `Under 55`

```js
assert(testLessThan(54) === 'Under 55');
```

`testLessThan(55)` يجب أن تُرجع السلسلة النصية `55 or Over`

```js
assert(testLessThan(55) === '55 or Over');
```

`testLessThan(99)` يجب أن تُرجع السلسلة النصية `55 or Over`

```js
assert(testLessThan(99) === '55 or Over');
```

يجب أن تستخدم معامل `<` مرتين على الأقل

```js
assert(__helpers.removeJSComments(code).match(/val\s*<\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testLessThan(val) {
  if (val) {  // Change this line
    return "Under 25";
  }

  if (val) {  // Change this line
    return "Under 55";
  }

  return "55 or Over";
}

testLessThan(10);
```

# --solutions--

```js
function testLessThan(val) {
  if (val < 25) {  // Change this line
    return "Under 25";
  }

  if (val < 55) {  // Change this line
    return "Under 55";
  }

  return "55 or Over";
}
```
