---
id: 56592a60ddddeae28f7aa8e1
title: الوصول إلى المصفوفات متعددة الأبعاد باستخدام الفهارس
challengeType: 1
forumTopicId: 16159
dashedName: access-multi-dimensional-arrays-with-indexes
---

# --description--

طريقة واحدة للتفكير في <dfn>المصفوفة متعددة الأبعاد</dfn> هي باعتبارها *مصفوفة من مصفوفات*. عند استخدام الأقواس للوصول إلى المصفوفة، تشير مجموعة الأقواس الأولى إلى المدخلات في المصفوفة الخارجية (المستوى الأول)، وكل زوج إضافي من الأقواس يشير إلى المستوى التالي من المدخلات داخلها.

**مثال**

```js
const arr = [
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9],
  [[10, 11, 12], 13, 14]
];

const subarray = arr[3];
const nestedSubarray = arr[3][0];
const element = arr[3][0][1];
```

في هذا المثال، تحتوي `subarray` على القيمة `[[10, 11, 12], 13, 14]`،
وتحتوي `nestedSubarray` على القيمة `[10, 11, 12]`، و`element` تحتوي على القيمة `11`.

**ملاحظة:** لا يجب وجود أي مسافات بين اسم المصفوفة والأقواس المربعة، مثل `array [0][0]` وحتى هذا `array [0] [0]` غير مسموح به. رغم أن JavaScript يستطيع معالجة هذا بشكل صحيح، قد يربك ذلك المبرمجين الآخرين الذين يقرؤون شفرتك.

# --instructions--

باستخدام ترميز الأقواس، اختر عنصرًا من `myArray` بحيث تكون قيمة `myData` مساوية لـ `8`.

# --hints--

يجب أن تكون قيمة `myData` مساوية لـ `8`.

```js
assert(myData === 8);
```

يجب استخدام ترميز الأقواس لقراءة القيمة الصحيحة من `myArray`.

```js
assert(/myData=myArray\[2\]\[1\]/.test(__helpers.removeWhiteSpace(__helpers.removeJSComments(code))));
```

# --seed--

## --seed-contents--

```js
const myArray = [
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9],
  [[10, 11, 12], 13, 14],
];

const myData = myArray[0][0];
```

# --solutions--

```js
const myArray = [[1, 2, 3], [4, 5, 6], [7, 8, 9], [[10, 11, 12], 13, 14]];
const myData = myArray[2][1];
```
