---
id: 565bbe00e9cc8ac0725390f4
title: عد البطاقات
challengeType: 1
forumTopicId: 16809
dashedName: counting-cards
---

# --description--

في لعبة الكازينو بلاك جاك، يمكن للاعب تحديد ما إذا كان لديه أفضلية في اليد التالية على الكازينو من خلال تتبع العدد النسبي للبطاقات العالية والمنخفضة المتبقية في الرزمة. يُسمى هذا عد البطاقات.

وجود المزيد من البطاقات العالية المتبقية في الرزمة يُفضل اللاعب. تُعيَّن قيمة لكل بطاقة وفقًا للجدول أدناه. عندما يكون العَد إيجابيًا، يجب على اللاعب المراهنة بمبلغ كبير. عندما يكون العَد صفرًا أو سلبيًا، يجب المراهنة بمبلغ صغير.

<table><thead><tr><th>تغير العَد</th><th>البطاقات</th></tr></thead><tbody><tr><td>+1</td><td>2, 3, 4, 5, 6</td></tr><tr><td>0</td><td>7, 8, 9</td></tr><tr><td>-1</td><td>10, 'J', 'Q', 'K', 'A'</td></tr></tbody></table>

ستكتب دالة عد البطاقات. ستستقبل معلمة `card` التي يمكن أن تكون رقمًا أو سلسلة نصية، وستزيد أو تنقص المتغير العالمي `count` حسب قيمة البطاقة (انظر الجدول). ثم تُرجع الدالة سلسلة نصية تحتوي على العَد الحالي وكلمة `Bet` إذا كان العَد إيجابيًا، أو `Hold` إذا كان العَد صفرًا أو سلبيًا. يجب أن يفصل بين العَد الحالي وقرار اللاعب (`Bet` أو `Hold`) مسافة واحدة فقط.

**أمثلة على المخرجات:** `-3 Hold` أو `5 Bet`

**تلميح**  
لا تُعدِّل قيمة `count` إلى 0 عندما تكون القيمة 7 أو 8 أو 9.  
لا تُرجع مصفوفة.  
لا تُدرج علامات اقتباس (مفردة أو مزدوجة) في المخرجات.

# --hints--

يجب أن تُرجع دالتك قيمة العَد والنص (`Bet` أو `Hold`) مع وجود مسافة واحدة بينهما.

```js
assert(//
  (function () {
    count = 0;
    let out = cc(10);
    const hasSpace = /-?\d+ (Bet|Hold)/.test('' + out);
    return hasSpace;
  })()
);
```

متتالية البطاقات 2، 3، 4، 5، 6 يجب أن تُرجع السلسلة النصية `5 Bet`

```js
assert(
  (function () {
    count = 0;
    cc(2);
    cc(3);
    cc(4);
    cc(5);
    var out = cc(6);
    if (out === '5 Bet') {
      return true;
    }
    return false;
  })()
);
```

متتالية البطاقات 7، 8، 9 يجب أن تُرجع السلسلة النصية `0 Hold`

```js
assert(
  (function () {
    count = 0;
    cc(7);
    cc(8);
    var out = cc(9);
    if (out === '0 Hold') {
      return true;
    }
    return false;
  })()
);
```

متتالية البطاقات 10، J، Q، K، A يجب أن تُرجع السلسلة النصية `-5 Hold`

```js
assert(
  (function () {
    count = 0;
    cc(10);
    cc('J');
    cc('Q');
    cc('K');
    var out = cc('A');
    if (out === '-5 Hold') {
      return true;
    }
    return false;
  })()
);
```

متتالية البطاقات 3، 7، Q، 8، A يجب أن تُرجع السلسلة النصية `-1 Hold`

```js
assert(
  (function () {
    count = 0;
    cc(3);
    cc(7);
    cc('Q');
    cc(8);
    var out = cc('A');
    if (out === '-1 Hold') {
      return true;
    }
    return false;
  })()
);
```

متتالية البطاقات 2، J، 9، 2، 7 يجب أن تُرجع السلسلة النصية `1 Bet`

```js
assert(
  (function () {
    count = 0;
    cc(2);
    cc('J');
    cc(9);
    cc(2);
    var out = cc(7);
    if (out === '1 Bet') {
      return true;
    }
    return false;
  })()
);
```

متتالية البطاقات 2، 2، 10 يجب أن تُرجع السلسلة النصية `1 Bet`

```js
assert(
  (function () {
    count = 0;
    cc(2);
    cc(2);
    var out = cc(10);
    if (out === '1 Bet') {
      return true;
    }
    return false;
  })()
);
```

متتالية البطاقات 3، 2، A، 10، K يجب أن تُرجع السلسلة النصية `-1 Hold`

```js
assert(
  (function () {
    count = 0;
    cc(3);
    cc(2);
    cc('A');
    cc(10);
    var out = cc('K');
    if (out === '-1 Hold') {
      return true;
    }
    return false;
  })()
);
```

# --seed--

## --seed-contents--

```js
let count = 0;

function cc(card) {
  // Only change code below this line


  return "Change Me";
  // Only change code above this line
}

cc(2); cc(3); cc(7); cc('K'); cc('A');
```

# --solutions--

```js
let count = 0;
function cc(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
