---
id: 5664820f61c48e80c9fa476c
title: شفرة الغولف
challengeType: 1
forumTopicId: 18195
dashedName: golf-code
---

# --description--

في لعبة الغولف، لكل حفرة قيمة `par`، وهي متوسط عدد `strokes` التي يُتوقع من لاعب الغولف أن يُنفذها لإدخال الكرة في الحفرة وإكمال اللعب. بناءً على مدى ارتفاع أو انخفاض `strokes` عن `par`، هناك لقب مختلف.

ستُمرر إلى دالتك المعلمتان `par` و`strokes`. أعد السلسلة النصية الصحيحة وفقًا لهذا الجدول الذي يسرد الخطوط حسب ترتيب الأولوية؛ من الأعلى (الأعلى أولوية) إلى الأسفل (الأقل أولوية):

<table><thead><tr><th>خطوط رسم</th><th>الإرجاع</th></tr></thead><tbody><tr><td>1</td><td>"Hole-in-one!"</td></tr><tr><td>&#x3C;= par - 2</td><td>"Eagle"</td></tr><tr><td>par - 1</td><td>"Birdie"</td></tr><tr><td>par</td><td>"Par"</td></tr><tr><td>par + 1</td><td>"Bogey"</td></tr><tr><td>par + 2</td><td>"Double Bogey"</td></tr><tr><td>>= par + 3</td><td>"Go Home!"</td></tr></tbody></table>

سيكون كل من `par` و`strokes` دائمًا رقميًا وموجبًا. أضفنا مصفوفة بكل الأسماء لملاءمتك.

# --hints--

`golfScore(4, 1)` يجب أن تُرجع السلسلة النصية `Hole-in-one!`

```js
assert(golfScore(4, 1) === 'Hole-in-one!');
```

`golfScore(4, 2)` يجب أن تُرجع السلسلة النصية `Eagle`

```js
assert(golfScore(4, 2) === 'Eagle');
```

`golfScore(5, 2)` يجب أن تُرجع السلسلة النصية `Eagle`

```js
assert(golfScore(5, 2) === 'Eagle');
```

`golfScore(4, 3)` يجب أن تُرجع السلسلة النصية `Birdie`

```js
assert(golfScore(4, 3) === 'Birdie');
```

`golfScore(4, 4)` يجب أن تُرجع السلسلة النصية `Par`

```js
assert(golfScore(4, 4) === 'Par');
```

`golfScore(1, 1)` يجب أن تُرجع السلسلة النصية `Hole-in-one!`

```js
assert(golfScore(1, 1) === 'Hole-in-one!');
```

`golfScore(5, 5)` يجب أن تُرجع السلسلة النصية `Par`

```js
assert(golfScore(5, 5) === 'Par');
```

`golfScore(4, 5)` يجب أن تُرجع السلسلة النصية `Bogey`

```js
assert(golfScore(4, 5) === 'Bogey');
```

`golfScore(4, 6)` يجب أن تُرجع السلسلة النصية `Double Bogey`

```js
assert(golfScore(4, 6) === 'Double Bogey');
```

`golfScore(4, 7)` يجب أن تُرجع السلسلة النصية `Go Home!`

```js
assert(golfScore(4, 7) === 'Go Home!');
```

`golfScore(5, 9)` يجب أن تُرجع السلسلة النصية `Go Home!`

```js
assert(golfScore(5, 9) === 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

function golfScore(par, strokes) {
  // Only change code below this line


  return "Change Me";
  // Only change code above this line
}

golfScore(5, 4);
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
