---
id: 5675e877dbd60be8ad28edc6
title: التكرار خلال مصفوفة باستخدام حلقة for
challengeType: 1
forumTopicId: 18216
dashedName: iterate-through-an-array-with-a-for-loop
---

# --description--

مهمة شائعة في JavaScript هي التكرار خلال محتويات مصفوفة. إحدى الطرق لذلك هي باستخدام حلقة `for`. ستُخرج هذه الشفرة كل عنصر من عناصر المصفوفة `arr` إلى وحدة التحكم:

```js
const arr = [10, 9, 8, 7, 6];

for (let i = 0; i < arr.length; i++) {
   console.log(arr[i]);
}
```

تذكر أن المصفوفات تبدأ من الصفر في الفهرسة، مما يعني أن الفهرس الأخير للمصفوفة هو `length - 1`. شرطنا لهذه الحلقة هو `i < arr.length`، الذي يوقف الحلقة عندما يصبح `i` مساويًا لـ `length`. في هذه الحالة، التكرار الأخير هو `i === 4` أي عندما يصبح `i` مساويًا لـ `arr.length - 1` ويُخرج `6` إلى وحدة التحكم. ثم يزيد `i` إلى `5`، وتنتهي الحلقة لأن `i < arr.length` تكون `false`.

# --instructions--

أعلن وابدأ متغيرًا باسم `total` بقيمة `0`. استخدم حلقة `for` لإضافة قيمة كل عنصر من مصفوفة `myArr` إلى `total`.

# --hints--

يجب إعلان وبدء `total` بقيمة 0.

```js
assert(__helpers.removeJSComments(code).match(/(var|let|const)\s*?total\s*=\s*0.*?;?/));
```

يجب أن يساوي `total` القيمة 20.

```js
assert(total === 20);
```

يجب استخدام حلقة `for` للتكرار خلال `myArr`.

```js
assert(/for\s*\(/g.test(__helpers.removeJSComments(code)) && /myArr\s*\[/g.test(__helpers.removeJSComments(code)));
```

لا تحاول تعيين القيمة 20 مباشرة إلى `total`.

```js
assert(!__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/total[=+-]0*[1-9]+/gm));
```

# --seed--

## --seed-contents--

```js
// Setup
const myArr = [2, 3, 4, 5, 6];

// Only change code below this line

```

# --solutions--

```js
const myArr = [2, 3, 4, 5, 6];
let total = 0;

for (let i = 0; i < myArr.length; i++) {
  total += myArr[i];
}
```
