---
id: 5679ceb97cbaa8c51670a16b
title: إرجاع قيم منطقية من الدوال
challengeType: 1
forumTopicId: 18273
dashedName: returning-boolean-values-from-functions
---

# --description--

قد تتذكر من <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/comparison-with-the-equality-operator" target="_blank" rel="noopener noreferrer nofollow">المقارنة باستخدام معامل المساواة</a> أن جميع معاملات المقارنة تُرجع قيمة منطقية `true` أو `false`.

أحيانًا يستخدم الناس عبارة `if/else` لإجراء مقارنة، مثل هذا:

```js
function isEqual(a, b) {
  if (a === b) {
    return true;
  } else {
    return false;
  }
}
```

لكن هناك طريقة أفضل للقيام بذلك. بما أن `===` تُرجع `true` أو `false`، يمكننا إرجاع نتيجة المقارنة مباشرة:

```js
function isEqual(a, b) {
  return a === b;
}
```

# --instructions--

صحح الدالة `isLess` لإزالة عبارات `if/else`.

# --hints--

يجب أن تُرجع `isLess(10, 15)` القيمة `true`

```js
assert(isLess(10, 15) === true);
```

يجب أن تُرجع `isLess(15, 10)` القيمة `false`

```js
assert(isLess(15, 10) === false);
```

يجب ألا تستخدم أي عبارات `if` أو `else`

```js
assert(!/if|else/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function isLess(a, b) {
  // Only change code below this line
  if (a < b) {
    return true;
  } else {
    return false;
  }
  // Only change code above this line
}

isLess(10, 15);
```

# --solutions--

```js
function isLess(a, b) {
  return a < b;
}
```
