---
id: 567af2437cbaa8c51670a16c
title: اختبار الكائنات للخصائص
challengeType: 1
forumTopicId: 18324
dashedName: testing-objects-for-properties
---

# --description--

للتحقق مما إذا كانت خاصية موجودة في كائن معين أم لا، يمكنك استخدام الطريقة `.hasOwnProperty()`. تعيد `someObject.hasOwnProperty(someProperty)` القيمة `true` أو `false` حسب ما إذا وُجدت الخاصية في الكائن أم لا.

**مثال**

```js
function checkForProperty(object, property) {
  return object.hasOwnProperty(property);
}

checkForProperty({ top: 'hat', bottom: 'pants' }, 'top'); // true
checkForProperty({ top: 'hat', bottom: 'pants' }, 'middle'); // false
```

يعيد الاستدعاء الأول للدالة `checkForProperty` القيمة `true`، بينما يعيد الاستدعاء الثاني القيمة `false`.

# --instructions--

عدّل الدالة `checkObj` لاختبار ما إذا كان الكائن الممرر إلى معلمة الدالة `obj` يحتوي الخاصية المحددة الممررة إلى معلمة الدالة `checkProp`. إذا وُجدت الخاصية الممررة إلى `checkProp` في `obj`، أعد قيمة تلك الخاصية. إذا لم توجد، أعد `Not Found`.

# --hints--

`checkObj({gift: "pony", pet: "kitten", bed: "sleigh"}, "gift")` يجب أن تعيد السلسلة النصية `pony`.

```js
assert(
  checkObj({ gift: 'pony', pet: 'kitten', bed: 'sleigh' }, 'gift') === 'pony'
);
```

`checkObj({gift: "pony", pet: "kitten", bed: "sleigh"}, "pet")` يجب أن تعيد السلسلة النصية `kitten`.

```js
assert(
  checkObj({ gift: 'pony', pet: 'kitten', bed: 'sleigh' }, 'pet') === 'kitten'
);
```

`checkObj({gift: "pony", pet: "kitten", bed: "sleigh"}, "house")` يجب أن تعيد السلسلة النصية `Not Found`.

```js
assert(
  checkObj({ gift: 'pony', pet: 'kitten', bed: 'sleigh' }, 'house') ===
    'Not Found'
);
```

`checkObj({city: "Seattle"}, "city")` يجب أن تعيد السلسلة النصية `Seattle`.

```js
assert(checkObj({ city: 'Seattle' }, 'city') === 'Seattle');
```

`checkObj({city: "Seattle"}, "district")` يجب أن تعيد السلسلة النصية `Not Found`.

```js
assert(checkObj({ city: 'Seattle' }, 'district') === 'Not Found');
```

`checkObj({pet: "kitten", bed: "sleigh"}, "gift")` يجب أن تعيد السلسلة النصية `Not Found`.

```js
assert(checkObj({ pet: 'kitten', bed: 'sleigh' }, 'gift') === 'Not Found');
```

# --seed--

## --seed-contents--

```js
function checkObj(obj, checkProp) {
  // Only change code below this line
  return "Change Me!";
  // Only change code above this line
}
```

# --solutions--

```js
function checkObj(obj, checkProp) {
  if(obj.hasOwnProperty(checkProp)) {
    return obj[checkProp];
  } else {
    return "Not Found";
  }
}
```
