---
id: 5688e62ea601b2482ff8422b
title: البحث في الملف الشخصي
challengeType: 1
forumTopicId: 18259
dashedName: profile-lookup
---

# --description--

لدينا مصفوفة من الكائنات تمثل أشخاصًا مختلفين في قوائم جهات الاتصال.

تم كتابة دالة `lookUpProfile` التي تأخذ `name` و`prop` كمعلمات مسبقًا من أجلك.

يجب أن تتحقق الدالة مما إذا كان `name` هو `firstName` لجهة اتصال فعلية، وإذا كانت الخاصية المعطاة (`prop`) هي خاصية لتلك الجهة.

إذا كان كلاهما صحيحًا، فارجع "قيمة" تلك الخاصية.

إذا لم يتطابق `name` مع أي جهة اتصال، فارجع السلسلة النصية `No such contact`.

إذا لم تتطابق `prop` مع أي خاصية صالحة لجهة اتصال تم العثور عليها تطابق `name`، فارجع السلسلة النصية `No such property`.

# --hints--

`lookUpProfile("Kristian", "lastName")` يجب أن تُرجع السلسلة النصية `Vos`

```js
assert(lookUpProfile('Kristian', 'lastName') === 'Vos');
```

`lookUpProfile("Sherlock", "likes")` يجب أن تُرجع `["Intriguing Cases", "Violin"]`

```js
assert.deepEqual(lookUpProfile('Sherlock', 'likes'), [
  'Intriguing Cases',
  'Violin'
]);
```

`lookUpProfile("Harry", "likes")` يجب أن تُرجع مصفوفة

```js
assert(typeof lookUpProfile('Harry', 'likes') === 'object');
```

`lookUpProfile("Bob", "number")` يجب أن تُرجع السلسلة النصية `No such contact`

```js
assert(lookUpProfile('Bob', 'number') === 'No such contact');
```

`lookUpProfile("Bob", "potato")` يجب أن تُرجع السلسلة النصية `No such contact`

```js
assert(lookUpProfile('Bob', 'potato') === 'No such contact');
```

`lookUpProfile("Akira", "address")` يجب أن تُرجع السلسلة النصية `No such property`

```js
assert(lookUpProfile('Akira', 'address') === 'No such property');
```

# --seed--

## --seed-contents--

```js
// Setup
const contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];

function lookUpProfile(name, prop) {
  // Only change code below this line

  // Only change code above this line
}

lookUpProfile("Akira", "likes");
```

# --solutions--

```js
const contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];
function lookUpProfile(name, prop) {
  for (let i in contacts) {
    if (contacts[i].firstName === name) {
      return contacts[i][prop] || "No such property";
    }
  }
  return "No such contact";
}
```
