---
id: 5690307fddb111c6084545d7
title: الترتيب المنطقي في عبارات If Else
challengeType: 1
forumTopicId: 18228
dashedName: logical-order-in-if-else-statements
---

# --description--

الترتيب مهم في عبارات `if` و `else if`.

تُنفذ الدالة من الأعلى إلى الأسفل، لذا يجب أن تكون حذرًا بشأن أي عبارة تأتي أولاً.

خذ هاتين الدالتين كمثال.

إليك الأولى:

```js
function foo(x) {
  if (x < 1) {
    return "Less than one";
  } else if (x < 2) {
    return "Less than two";
  } else {
    return "Greater than or equal to two";
  }
}
```

والثانية فقط تبدل ترتيب العبارات:

```js
function bar(x) {
  if (x < 2) {
    return "Less than two";
  } else if (x < 1) {
    return "Less than one";
  } else {
    return "Greater than or equal to two";
  }
}
```

بينما تبدو هاتان الدالتان متطابقتين تقريبًا، إذا مررنا رقمًا إلى كلتيهما نحصل على مخرجات مختلفة.

```js
foo(0)
bar(0)
```

`foo(0)` سترجع السلسلة النصية `Less than one`، و `bar(0)` سترجع السلسلة النصية `Less than two`.

# --instructions--

غيّر ترتيب المنطق في الدالة بحيث تُرجع العبارات الصحيحة في جميع الحالات.

# --hints--

`orderMyLogic(4)` يجب أن تُرجع السلسلة النصية `Less than 5`

```js
assert(orderMyLogic(4) === 'Less than 5');
```

`orderMyLogic(6)` يجب أن تُرجع السلسلة النصية `Less than 10`

```js
assert(orderMyLogic(6) === 'Less than 10');
```

`orderMyLogic(11)` يجب أن تُرجع السلسلة النصية `Greater than or equal to 10`

```js
assert(orderMyLogic(11) === 'Greater than or equal to 10');
```

# --seed--

## --seed-contents--

```js
function orderMyLogic(val) {
  if (val < 10) {
    return "Less than 10";
  } else if (val < 5) {
    return "Less than 5";
  } else {
    return "Greater than or equal to 10";
  }
}

orderMyLogic(7);
```

# --solutions--

```js
function orderMyLogic(val) {
  if(val < 5) {
    return "Less than 5";
  } else if (val < 10) {
    return "Less than 10";
  } else {
    return "Greater than or equal to 10";
  }
}
```
