---
id: 56bbb991ad1ed5201cd392cf
title: كتابة جافاسكريبت قابلة لإعادة الاستخدام باستخدام الدوال
challengeType: 1
forumTopicId: 18378
dashedName: write-reusable-javascript-with-functions
---

# --description--

في جافاسكريبت، يمكننا تقسيم الشفرة إلى أجزاء قابلة لإعادة الاستخدام تُسمى <dfn>دوال</dfn>.

إليك مثالًا على دالة:

```js
function functionName() {
  console.log("Hello World");
}
```

يمكنك استدعاء أو <dfn>تنفيذ</dfn> هذه الدالة باستخدام اسمها متبوعًا بالأقواس العادية، مثل هذا: `functionName();` في كل مرة تُستدعى فيها الدالة، ستطبع الرسالة `Hello World` في وحدة التحكم. كل الشفرة بين الأقواس المتعرجة تُنفذ في كل استدعاء للدالة.

# --instructions--

<ol>
  <li>
أنشئ دالة باسم <code>reusableFunction</code> تطبع السلسلة النصية <code>Hi World</code> في وحدة التحكم.
  </li>
  <li>
استدعِ الدالة.
  </li>
</ol>

# --before-each--

```js
function testConsole() {
  var originalConsole = console;
  var nativeLog = console.log;
  var hiWorldWasLogged = false;
  console.log = function (message) {
    if (message === 'Hi World') {
      console.warn(message);
      hiWorldWasLogged = true;
    }
    if (nativeLog.apply) {
      nativeLog.apply(originalConsole, arguments);
    } else {
      var nativeMsg = Array.prototype.slice.apply(arguments).join(' ');
      nativeLog(nativeMsg);
    }
  };
  reusableFunction();
  console.log = nativeLog;
  return hiWorldWasLogged;
}
```

# --hints--

`reusableFunction` يجب أن تكون دالة.

```js
assert(typeof reusableFunction === 'function');
```

إذا تم استدعاء `reusableFunction`، يجب أن تُخرج السلسلة النصية `Hi World` في وحدة التحكم.

```js
assert(testConsole());
```

يجب استدعاء `reusableFunction` بعد تعريفها.

```js
const functionStr = reusableFunction && __helpers.removeWhiteSpace(reusableFunction.toString());
const codeWithoutFunction = __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).replace(/reusableFunction\(\)\{/g, '');
assert(/reusableFunction\(\)/.test(codeWithoutFunction));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reusableFunction() {
  console.log("Hi World");
}
reusableFunction();
```
