---
id: 56bbb991ad1ed5201cd392cc
title: التلاعب بالمصفوفات باستخدام طريقة pop
challengeType: 1
forumTopicId: 18236
dashedName: manipulate-arrays-with-pop
---

# --description--

طريقة أخرى لتغيير البيانات في المصفوفة هي باستخدام الدالة `.pop()`.

تُستخدم `.pop()` لإزالة قيمة من نهاية المصفوفة. يمكننا تخزين هذه القيمة التي أُزيلت بتعيينها إلى متغير. بعبارة أخرى، تُزيل `.pop()` العنصر الأخير من المصفوفة وتُعيد هذا العنصر.

يمكن إزالة أي نوع من العناصر من المصفوفة - أرقام، سلاسل نصية، وحتى مصفوفات مسكنة.

```js
const threeArr = [1, 4, 6];
const oneDown = threeArr.pop();
console.log(oneDown);
console.log(threeArr);
```

سيعرض أول `console.log` القيمة `6`، وسيعرض الثاني القيمة `[1, 4]`.

# --instructions--

استخدم الدالة `.pop()` لإزالة العنصر الأخير من `myArray` وعيّن القيمة التي أُزيلت إلى متغير جديد، `removedFromMyArray`.

# --hints--

يجب أن يحتوي `myArray` فقط على `[["John", 23]]`.

```js
assert(
  (function (d) {
    if (d[0][0] == 'John' && d[0][1] === 23 && d[1] == undefined) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

يجب استخدام `pop()` على `myArray`.

```js
assert(/removedFromMyArray\s*=\s*myArray\s*.\s*pop\s*(\s*)/.test(__helpers.removeJSComments(code)));
```

يجب أن يحتوي `removedFromMyArray` فقط على `["cat", 2]`.

```js
assert(
  (function (d) {
    if (d[0] == 'cat' && d[1] === 2 && d[2] == undefined) {
      return true;
    } else {
      return false;
    }
  })(removedFromMyArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["cat", 2]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["cat", 2]];
const removedFromMyArray = myArray.pop();
```
