---
id: 56bbb991ad1ed5201cd392ca
title: الوصول إلى بيانات المصفوفة باستخدام الفهارس
challengeType: 1
forumTopicId: 16158
dashedName: access-array-data-with-indexes
---

# --description--

يمكننا الوصول إلى البيانات داخل المصفوفات باستخدام <dfn>الفهارس</dfn>.

تُكتب فهارس المصفوفة بنفس ترميز الأقواس الذي تستخدمه السلاسل النصية، لكن بدلًا من تحديد حرف، تحدد عنصرًا في المصفوفة. مثل السلاسل النصية، تستخدم المصفوفات ترقيمًا <dfn>يبدأ من الصفر</dfn>، لذا العنصر الأول في المصفوفة له فهرس `0`.

<br>

**مثال**

```js
const array = [50, 60, 70];
console.log(array[0]);
const data = array[1];
```

تطبع `console.log(array[0])` القيمة `50`، وتكون قيمة `data` هي `60`.

# --instructions--

أنشئ متغيرًا اسمه `myData` واجعله يساوي القيمة الأولى في `myArray` باستخدام ترميز الأقواس.

# --hints--

يجب أن يساوي المتغير `myData` القيمة الأولى في `myArray`.

```js
assert(
  (function () {
    if (
      typeof myArray !== 'undefined' &&
      typeof myData !== 'undefined' &&
      myArray[0] === myData
    ) {
      return true;
    } else {
      return false;
    }
  })()
);
```

يجب الوصول إلى البيانات في المتغير `myArray` باستخدام ترميز الأقواس.

```js
assert(
  (function () {
    if (__helpers.removeJSComments(code).match(/\s*=\s*myArray\[0\]/g)) {
      return true;
    } else {
      return false;
    }
  })()
);
```

# --seed--

## --seed-contents--

```js
const myArray = [50, 60, 70];


```

# --solutions--

```js
const myArray = [50, 60, 70];
const myData = myArray[0];
```
