---
id: 56bbb991ad1ed5201cd392ce
title: التلاعب بالمصفوفات باستخدام طريقة unshift
challengeType: 1
forumTopicId: 18239
dashedName: manipulate-arrays-with-unshift
---

# --description--

لا يمكنك فقط استخدام `shift` لإزالة عناصر من بداية المصفوفة، بل يمكنك أيضًا استخدام `unshift` لإضافة عناصر في بداية المصفوفة، أي إدخال عناصر أمام المصفوفة.

تعمل `.unshift()` تمامًا مثل `.push()`، لكن بدلًا من إضافة العنصر في نهاية المصفوفة، تضيف `unshift()` العنصر في بداية المصفوفة.

مثال:

```js
const ourArray = ["Stimpson", "J", "cat"];
ourArray.shift();
ourArray.unshift("Happy");
```

بعد استخدام `shift`، ستكون قيمة `ourArray` هي `["J", "cat"]`. وبعد استخدام `unshift`، ستكون قيمة `ourArray` هي `["Happy", "J", "cat"]`.

# --instructions--

أضف `["Paul", 35]` إلى بداية المتغير `myArray` باستخدام `unshift()`.

# --hints--

يجب أن تحتوي `myArray` الآن على `[["Paul", 35], ["dog", 3]]`.

```js
assert(
  (function (d) {
    if (
      typeof d[0] === 'object' &&
      d[0][0] == 'Paul' &&
      d[0][1] === 35 &&
      d[1][0] != undefined &&
      d[1][0] == 'dog' &&
      d[1][1] != undefined &&
      d[1][1] == 3
    ) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["dog", 3]];
myArray.shift();

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["dog", 3]];
myArray.shift();
myArray.unshift(["Paul", 35]);
```
