---
id: 56bbb991ad1ed5201cd392d1
title: تحديث خصائص الكائن
challengeType: 1
forumTopicId: 18336
dashedName: updating-object-properties
---

# --description--

بعد إنشاء كائن JavaScript، يمكنك تحديث خصائصه في أي وقت تمامًا كما تحدث أي متغير آخر. يمكنك استخدام ترميز النقطة أو ترميز الأقواس للتحديث.

على سبيل المثال، لننظر إلى `ourDog`:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```

بما أنه كلب سعيد جدًا، دعنا نغير اسمه إلى السلسلة النصية `Happy Camper`. هكذا نُحدث خاصية الاسم في كائنه: `ourDog.name = "Happy Camper";` أو `ourDog["name"] = "Happy Camper";` الآن عند تقييم `ourDog.name`، بدلًا من الحصول على `Camper`، سنحصل على اسمه الجديد، `Happy Camper`.

# --instructions--

حدّث خاصية الاسم في كائن `myDog`. دعنا نغير اسمها من `Coder` إلى `Happy Coder`. يمكنك استخدام ترميز النقطة أو ترميز الأقواس.

# --hints--

يجب أن تحدّث خاصية `name` في `myDog` لتساوي السلسلة النصية `Happy Coder`.

```js
assert(/happy coder/gi.test(myDog.name));
```

لا يجب تعديل تعريف `myDog`.

```js
assert(/"name": "Coder"/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const myDog = {
  "name": "Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};

// Only change code below this line

```

# --solutions--

```js
const myDog = {
  "name": "Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};
myDog.name = "Happy Coder";
```
