---
id: 56bbb991ad1ed5201cd392d2
title: إضافة خصائص جديدة إلى كائن JavaScript
challengeType: 1
forumTopicId: 301169
dashedName: add-new-properties-to-a-javascript-object
---

# --description--

يمكنك إضافة خصائص جديدة إلى كائنات JavaScript الموجودة بنفس الطريقة التي تعدلها بها.

إليك كيف نضيف خاصية `bark` إلى `ourDog`:

```js
ourDog.bark = "bow-wow";
```

أو

```js
ourDog["bark"] = "bow-wow";
```

الآن عندما نقيم `ourDog.bark`، سنحصل على نباحه، `bow-wow`.

مثال:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};

ourDog.bark = "bow-wow";
```

# --instructions--

أضف خاصية `bark` إلى `myDog` واجعلها صوت كلب، مثل "woof". يمكنك استخدام ترميز النقطة أو ترميز الأقواس.

# --hints--

يجب أن تضيف الخاصية `bark` إلى `myDog`.

```js
assert(myDog.bark !== undefined);
```

لا يجب أن تضيف `bark` إلى تهيئة `myDog`.

```js
assert(!/bark[^\n]:/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};


```

# --solutions--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};
myDog.bark = "Woof Woof";
```
