---
id: 56bbb991ad1ed5201cd392cb
title: التلاعب بالمصفوفات باستخدام طريقة push
challengeType: 1
forumTopicId: 18237
dashedName: manipulate-arrays-with-push
---

# --description--

طريقة سهلة لإضافة بيانات إلى نهاية المصفوفة هي عبر الطريقة `push()`.

تأخذ الطريقة `push()` معلمة واحدة أو أكثر وتضيفها إلى نهاية المصفوفة، بالترتيب الذي تظهر به. تُرجع الطول الجديد للمصفوفة.

أمثلة:

```js
const arr1 = [1, 2, 3];
arr1.push(4, 5);

const arr2 = ["Stimpson", "J", "cat"];
arr2.push(["happy", "joy"]);
```

القيمة الحالية لـ `arr1` هي `[1, 2, 3, 4, 5]` والقيمة الحالية لـ `arr2` هي `["Stimpson", "J", "cat", ["happy", "joy"]]`.

# --instructions--

أضف `["dog", 3]` إلى نهاية المتغير `myArray`.

# --hints--

يجب أن تساوي `myArray` الآن `[["John", 23], ["cat", 2], ["dog", 3]]`.

```js
assert(
  (function (d) {
    if (
      d[2] != undefined &&
      d[0][0] == 'John' &&
      d[0][1] === 23 &&
      d[2][0] == 'dog' &&
      d[2][1] === 3 &&
      d[2].length == 2
    ) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["cat", 2]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["cat", 2]];
myArray.push(["dog",3]);
```
