---
id: 56bbb991ad1ed5201cd392d0
title: بناء كائنات JavaScript
challengeType: 1
forumTopicId: 16769
dashedName: build-javascript-objects
---

# --description--

ربما سمعت مصطلح `object` من قبل.

الكائنات تشبه `arrays`، لكن بدلًا من استخدام الفهارس للوصول إلى بياناتها وتعديلها، تصل إلى البيانات في الكائنات عبر ما يسمى بـ `properties`.

الكائنات مفيدة لتخزين البيانات بطريقة منظمة، ويمكنها تمثيل أشياء من العالم الحقيقي، مثل القط.

إليك مثال على كائن قط:

```js
const cat = {
  "name": "Whiskers",
  "legs": 4,
  "tails": 1,
  "enemies": ["Water", "Dogs"]
};
```

في هذا المثال، تُخزن كل الخصائص كسلاسل نصية، مثل `name` و`legs` و`tails`. مع ذلك، يمكنك أيضًا استخدام الأرقام كخصائص. يمكنك حتى حذف علامات الاقتباس للخصائص النصية ذات الكلمة الواحدة، كما يلي:

```js
const anotherObject = {
  make: "Ford",
  5: "five",
  "model": "focus"
};
```

لكن إذا كان للكائن الخاص بك أي خصائص ليست نصية، فسيقوم JavaScript تلقائيًا بـ تحويل النوع إلى سلسلة نصية.

# --instructions--

اصنع كائنًا يمثل كلبًا اسمه `myDog` يحتوي على الخصائص `name` (سلسلة نصية)، و`legs`، و`tails`، و`friends`.

يمكنك تعيين قيم هذه الخصائص كما تريد، طالما أن `name` سلسلة نصية، و`legs` و`tails` أرقام، و`friends` مصفوفة.

# --hints--

يجب أن يحتوي `myDog` على الخاصية `name` ويجب أن تكون `string`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('name') &&
      z.name !== undefined &&
      typeof z.name === 'string'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

يجب أن يحتوي `myDog` على الخاصية `legs` ويجب أن تكون `number`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('legs') &&
      z.legs !== undefined &&
      typeof z.legs === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

يجب أن يحتوي `myDog` على الخاصية `tails` ويجب أن تكون `number`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('tails') &&
      z.tails !== undefined &&
      typeof z.tails === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

يجب أن يحتوي `myDog` على الخاصية `friends` ويجب أن تكون `array`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('friends') &&
      z.friends !== undefined &&
      Array.isArray(z.friends)
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

يجب أن يحتوي `myDog` فقط على كل الخصائص المعطاة.

```js
assert(
  (function (z) {
    return Object.keys(z).length === 4;
  })(myDog)
);
```

# --seed--

## --seed-contents--

```js
const myDog = {
  // Only change code below this line


  // Only change code above this line
};
```

# --solutions--

```js
const myDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```
