---
id: 56bbb991ad1ed5201cd392cd
title: التلاعب بالمصفوفات باستخدام طريقة shift
challengeType: 1
forumTopicId: 18238
dashedName: manipulate-arrays-with-shift
---

# --description--

دالة `pop()` تزيل دائمًا العنصر الأخير من المصفوفة. ماذا لو أردت إزالة العنصر الأول؟

هنا تأتي دالة `.shift()`. تعمل تمامًا مثل `.pop()`، لكنها تزيل العنصر الأول بدلًا من الأخير.

مثال:

```js
const ourArray = ["Stimpson", "J", ["cat"]];
const removedFromOurArray = ourArray.shift();
```

سيكون لـ `removedFromOurArray` قيمة السلسلة النصية `Stimpson`، وستصبح `ourArray` تساوي `["J", ["cat"]]`.

# --instructions--

استخدم دالة `.shift()` لإزالة العنصر الأول من `myArray` وعيّن القيمة "المزالة" إلى متغير جديد، `removedFromMyArray`.

# --hints--

يجب أن تساوي `myArray` الآن `[["dog", 3]]`.

```js
assert(
  (function (d) {
    if (d[0][0] == 'dog' && d[0][1] === 3 && d[1] == undefined) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

يجب أن يحتوي `removedFromMyArray` على `["John", 23]`.

```js
assert(
  (function (d) {
    if (
      d[0] == 'John' &&
      d[1] === 23 &&
      typeof removedFromMyArray === 'object'
    ) {
      return true;
    } else {
      return false;
    }
  })(removedFromMyArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["dog", 3]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["dog", 3]];

// Only change code below this line
const removedFromMyArray = myArray.shift();
```
