---
id: 56bbb991ad1ed5201cd392d3
title: حذف الخصائص من كائن JavaScript
challengeType: 1
forumTopicId: 17560
dashedName: delete-properties-from-a-javascript-object
---

# --description--

يمكننا أيضًا حذف الخصائص من الكائنات بهذه الطريقة:

```js
delete ourDog.bark;
```

مثال:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"],
  "bark": "bow-wow"
};

delete ourDog.bark;
```

بعد السطر الأخير الموضح أعلاه، يبدو `ourDog` كما يلي:

```js
{
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
}
```

# --instructions--

احذف الخاصية `tails` من `myDog`. يمكنك استخدام ترميز النقطة أو ترميز الأقواس.

# --hints--

يجب حذف الخاصية `tails` من `myDog`.

```js
assert(typeof myDog === 'object' && myDog.tails === undefined);
```

لا يجب تعديل إعداد `myDog`.

```js
assert(__helpers.removeJSComments(code).match(/"tails": 1/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"],
  "bark": "woof"
};

// Only change code below this line

```

# --solutions--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"],
  "bark": "woof"
};
delete myDog.tails;
```
