---
id: 587d7b7e367417b2b2512b24
title: استخدم المعامل الشرطي (الثلاثي)
challengeType: 1
forumTopicId: 301181
dashedName: use-the-conditional-ternary-operator
---

# --description--

يمكن استخدام <dfn>المعامل الشرطي</dfn>، المعروف أيضًا باسم <dfn>المعامل الثلاثي</dfn>، كتعبير if-else في سطر واحد.

الصيغة النحوية هي `a ? b : c`، حيث `a` هو الشرط، و`b` هو الشفرة التي تُنفذ عندما يُرجع الشرط `true`، و`c` هو الشفرة التي تُنفذ عندما يُرجع الشرط `false`.

تستخدم الدالة التالية عبارة `if/else` لفحص شرط:

```js
function findGreater(a, b) {
  if(a > b) {
    return "a is greater";
  }
  else {
    return "b is greater or equal";
  }
}
```

يمكن إعادة كتابة ذلك باستخدام المعامل الشرطي:

```js
function findGreater(a, b) {
  return a > b ? "a is greater" : "b is greater or equal";
}
```

# --instructions--

استخدم المعامل الشرطي في الدالة `checkEqual` للتحقق مما إذا كان رقمان متساويين أم لا. يجب أن تُرجع الدالة إما السلسلة النصية `Equal` أو السلسلة النصية `Not Equal`.

# --hints--

يجب أن تستخدم `checkEqual` المعامل الشرطي

```js
assert(/.+?\s*?\?\s*?.+?\s*?:\s*?.+?/.test(__helpers.removeJSComments(code)));
```

يجب أن تُرجع `checkEqual(1, 2)` السلسلة النصية `Not Equal`

```js
assert(checkEqual(1, 2) === 'Not Equal');
```

يجب أن تُرجع `checkEqual(1, 1)` السلسلة النصية `Equal`

```js
assert(checkEqual(1, 1) === 'Equal');
```

يجب أن تُرجع `checkEqual(1, -1)` السلسلة النصية `Not Equal`

```js
assert(checkEqual(1, -1) === 'Not Equal');
```

# --seed--

## --seed-contents--

```js
function checkEqual(a, b) {

}

checkEqual(1, 2);
```

# --solutions--

```js
function checkEqual(a, b) {
  return a === b ? "Equal" : "Not Equal";
}
```
