---
id: 587d7b7e367417b2b2512b21
title: استخدام عدة عوامل شرطية (ثلاثية)
challengeType: 1
forumTopicId: 301179
dashedName: use-multiple-conditional-ternary-operators
---

# --description--

في التحدي السابق، استخدمت عاملًا شرطيًا واحدًا. يمكنك أيضًا ربطها معًا للتحقق من شروط متعددة.

تستخدم الدالة التالية عبارات `if` و `else if` و `else` للتحقق من شروط متعددة:

```js
function findGreaterOrEqual(a, b) {
  if (a === b) {
    return "a and b are equal";
  }
  else if (a > b) {
    return "a is greater";
  }
  else {
    return "b is greater";
  }
}
```

يمكن إعادة كتابة الدالة أعلاه باستخدام عدة عوامل شرطية:

```js
function findGreaterOrEqual(a, b) {
  return (a === b) ? "a and b are equal" 
    : (a > b) ? "a is greater" 
    : "b is greater";
}
```

يُعتبر من أفضل الممارسات تنسيق عدة عوامل شرطية بحيث يكون كل شرط في سطر منفصل، كما هو موضح أعلاه. استخدام عدة عوامل شرطية بدون تنسيق مناسب قد يجعل شفرتك صعبة القراءة. على سبيل المثال:

```js
function findGreaterOrEqual(a, b) {
  return (a === b) ? "a and b are equal" : (a > b) ? "a is greater" : "b is greater";
}
```

# --instructions--

في الدالة `checkSign`، استخدم عدة عوامل شرطية - باتباع التنسيق الموصى به المستخدم في `findGreaterOrEqual` - للتحقق مما إذا كان الرقم موجبًا أو سالبًا أو صفرًا. يجب أن تُرجع الدالة `positive` أو `negative` أو `zero`.

# --hints--

يجب أن تستخدم `checkSign` عدة عوامل شرطية

```js
assert(/.+?\s*?\?\s*?.+?\s*?:\s*?.+?\s*?\?\s*?.+?\s*?:\s*?.+?/gi.test(__helpers.removeJSComments(code)));
```

يجب أن تُرجع `checkSign(10)` السلسلة النصية `positive`. لاحظ أن حالة الأحرف مهمة

```js
assert(checkSign(10) === 'positive');
```

يجب أن تُرجع `checkSign(-12)` السلسلة النصية `negative`. لاحظ أن حالة الأحرف مهمة

```js
assert(checkSign(-12) === 'negative');
```

يجب أن تُرجع `checkSign(0)` السلسلة النصية `zero`. لاحظ أن حالة الأحرف مهمة

```js
assert(checkSign(0) === 'zero');
```

# --seed--

## --seed-contents--

```js
function checkSign(num) {

}

checkSign(10);
```

# --solutions--

```js
function checkSign(num) {
  return (num > 0) ? 'positive' : (num < 0) ? 'negative' : 'zero';
}
```
