---
id: 587d7b7e367417b2b2512b23
title: استخدام الدالة parseInt
challengeType: 1
forumTopicId: 301183
dashedName: use-the-parseint-function
---

# --description--

الدالة `parseInt()` تُحلل سلسلة نصية وتُرجع عددًا صحيحًا. إليك مثالًا:

```js
const a = parseInt("007");
```

الدالة أعلاه تُحوّل السلسلة النصية `007` إلى العدد الصحيح `7`. إذا لم يكن بالإمكان تحويل الحرف الأول في السلسلة إلى رقم، فترجع `NaN`.

# --instructions--

استخدم `parseInt()` في الدالة `convertToInteger` لتحويل السلسلة النصية المُدخلة `str` إلى عدد صحيح، ثم إرجاعه.

# --hints--

يجب أن تستخدم `convertToInteger` الدالة `parseInt()`

```js
assert(/parseInt/g.test(__helpers.removeJSComments(code)));
```

يجب أن تُرجع `convertToInteger("56")` عددًا

```js
assert(typeof convertToInteger('56') === 'number');
```

يجب أن تُرجع `convertToInteger("56")` القيمة 56

```js
assert(convertToInteger('56') === 56);
```

يجب أن تُرجع `convertToInteger("77")` القيمة 77

```js
assert(convertToInteger('77') === 77);
```

يجب أن تُرجع `convertToInteger("JamesBond")` القيمة `NaN`

```js
assert.isNaN(convertToInteger('JamesBond'));
```

# --seed--

## --seed-contents--

```js
function convertToInteger(str) {

}

convertToInteger("56");
```

# --solutions--

```js
function convertToInteger(str) {
  return parseInt(str);
}
```
