---
id: 587d7b7e367417b2b2512b22
title: استخدام الدالة parseInt مع أساس العد
challengeType: 1
forumTopicId: 301182
dashedName: use-the-parseint-function-with-a-radix
---

# --description--

تُحلل الدالة `parseInt()` سلسلة نصية وتُرجع عددًا صحيحًا. تأخذ معلمة ثانية لأساس العد، الذي يحدد أساس الرقم في السلسلة النصية. يمكن أن يكون أساس العد عددًا صحيحًا بين 2 و36.

يبدو استدعاء الدالة كالتالي:

```js
parseInt(string, radix);
```

وهنا مثال:

```js
const a = parseInt("11", 2);
```

تُشير متغير أساس العد إلى أن `11` في النظام الثنائي، أو أساس العد 2. يحول هذا المثال السلسلة النصية `11` إلى العدد الصحيح `3`.

# --instructions--

استخدم `parseInt()` في الدالة `convertToInteger` لتحويل رقم ثنائي إلى عدد صحيح وإرجاعه.

# --hints--

يجب أن تستخدم `convertToInteger` الدالة `parseInt()`

```js
assert(/parseInt/g.test(__helpers.removeJSComments(code)));
```

يجب أن تُرجع `convertToInteger("10011")` عددًا

```js
assert(typeof convertToInteger('10011') === 'number');
```

يجب أن تُرجع `convertToInteger("10011")` العدد 19

```js
assert(convertToInteger('10011') === 19);
```

يجب أن تُرجع `convertToInteger("111001")` العدد 57

```js
assert(convertToInteger('111001') === 57);
```

يجب أن تُرجع `convertToInteger("JamesBond")` القيمة `NaN`

```js
assert.isNaN(convertToInteger('JamesBond'));
```

# --seed--

## --seed-contents--

```js
function convertToInteger(str) {

}

convertToInteger("10011");
```

# --solutions--

```js
function convertToInteger(str) {
  return parseInt(str, 2);
}
```
