---
id: 587d7b87367417b2b2512b3f
title: استكشاف الفروقات بين الكلمات المفتاحية var و let
challengeType: 1
forumTopicId: 301202
dashedName: explore-differences-between-the-var-and-let-keywords
---

# --description--

واحدة من أكبر المشاكل عند إعلان المتغيرات باستخدام الكلمة المفتاحية `var` هي أنه يمكنك بسهولة الكتابة فوق إعلانات المتغيرات:

```js
var camper = "James";
var camper = "David";
console.log(camper);
```

في الشفرة أعلاه، تم إعلان المتغير `camper` في الأصل كـ `James`، ثم تم استبداله ليصبح `David`. تعرض وحدة التحكم بعدها السلسلة النصية `David`.

في تطبيق صغير، قد لا تواجه هذا النوع من المشاكل. لكن مع زيادة حجم قاعدة الشفرة، قد تكتب فوق متغير لم تقصد تغييره. وبما أن هذا السلوك لا يرمي خطأً برمجيًا، يصبح البحث عن الأخطاء البرمجية وإصلاحها أكثر صعوبة.

تم تقديم كلمة مفتاحية تسمى `let` في ES6، وهو تحديث رئيسي لجافاسكريبت، لحل هذه المشكلة المحتملة مع الكلمة المفتاحية `var`. ستتعلم عن ميزات أخرى في ES6 في تحديات لاحقة.

إذا استبدلت `var` بـ `let` في الشفرة أعلاه، سينتج عن ذلك خطأ:

```js
let camper = "James";
let camper = "David";
```

يمكن رؤية الخطأ في وحدة التحكم في المتصفح.

لذا، على عكس `var`، عند استخدام `let`، يمكن إعلان متغير بنفس الاسم مرة واحدة فقط.

# --instructions--

حدّث الشفرة بحيث تستخدم فقط الكلمة المفتاحية `let`.

# --hints--

يجب ألا توجد كلمة `var` في الشفرة.

```js
assert.notMatch(code, /var/g);
```

يجب أن تكون قيمة `catName` هي السلسلة النصية `Oliver`.

```js
assert.equal(catName, 'Oliver');
```

يجب أن تكون قيمة `catSound` هي السلسلة النصية `Meow!`.

```js
assert.equal(catSound, 'Meow!');
```

# --seed--

## --seed-contents--

```js
var catName = "Oliver";
var catSound = "Meow!";
```

# --solutions--

```js
let catName = "Oliver";
let catSound = "Meow!";
```
