---
id: 598e8944f009e646fc236146
title: فهم القيمة غير المعرفة التي تُرجع من دالة
challengeType: 1
forumTopicId: 301177
dashedName: understanding-undefined-value-returned-from-a-function
---

# --description--

يمكن أن تتضمن الدالة جملة `return` لكنها ليست ملزمة بذلك. إذا لم تحتوي الدالة على جملة `return`، عند استدعائها، تنفذ الدالة الشفرة الداخلية لكن القيمة المرجعة تكون `undefined`.

**مثال**

```js
let sum = 0;

function addSum(num) {
  sum = sum + num;
}

addSum(3);
```

`addSum` دالة لا تحتوي على جملة `return`. ستغير الدالة المتغير العالمي `sum` لكن القيمة المرجعة من الدالة هي `undefined`.

# --instructions--

أنشئ دالة `addFive` بدون معلمات. تضيف هذه الدالة 5 إلى المتغير `sum`، لكن القيمة المرجعة منها تكون `undefined`.

# --hints--

يجب أن تكون `addFive` دالة.

```js
assert(typeof addFive === 'function');
```

بعد تشغيل الدالتين، يجب أن تكون قيمة `sum` مساوية لـ `8`.

```js
assert(sum === 8);
```

القيمة المرجعة من `addFive` يجب أن تكون `undefined`.

```js
assert(addFive() === undefined);
```

داخل الدالة `addFive`، يجب أن تضيف `5` إلى المتغير `sum`.

```js
assert(
  __helpers.removeWhiteSpace(addFive.toString()).match(/sum=sum\+5|sum\+=5/)
);
```

# --seed--

## --seed-contents--

```js
// Setup
let sum = 0;

function addThree() {
  sum = sum + 3;
}

// Only change code below this line


// Only change code above this line

addThree();
addFive();
```

# --solutions--

```js
let sum = 0;

function addThree() {
  sum = sum + 3;
}

function addFive() {
  sum = sum + 5;
}

addThree();
addFive();
```
