---
id: 599a789b454f2bbd91a3ff4d
title: ممارسة مقارنة القيم المختلفة
challengeType: 1
forumTopicId: 301174
dashedName: practice-comparing-different-values
---

# --description--

في التحديين الأخيرين، تعلّمنا عن معامل المساواة (`==`) ومعامل المساواة الصارمة (`===`). لنقم بمراجعة سريعة وممارسة استخدام هذين المعاملين أكثر.

إذا لم تكن القيم التي تتم مقارنتها من نفس النوع، يقوم معامل المساواة بتحويل النوع ثم تقييم القيم. أما معامل المساواة الصارمة فيقارن نوع البيانات والقيمة كما هي، دون تحويل أي نوع إلى الآخر.

**أمثلة**

`3 == '3'` تُرجع `true` لأن JavaScript يقوم بتحويل النوع من سلسلة نصية إلى رقم. `3 === '3'` تُرجع `false` لأن الأنواع مختلفة ولا يتم تحويل النوع.

**ملاحظة:** في JavaScript، يمكنك تحديد نوع متغير أو قيمة باستخدام معامل `typeof`، كما يلي:

```js
typeof 3
typeof '3'
```

`typeof 3` تُرجع السلسلة النصية `number`، و`typeof '3'` تُرجع السلسلة النصية `string`.

# --instructions--

الدالة `compareEquality` في المحرر تقارن بين قيمتين باستخدام معامل المساواة. عدّل الدالة بحيث تُرجع السلسلة النصية `Equal` فقط عندما تكون القيم متساوية صارمًا.

# --hints--

`compareEquality(10, "10")` يجب أن تُرجع السلسلة النصية `Not Equal`

```js
assert(compareEquality(10, '10') === 'Not Equal');
```

`compareEquality("20", 20)` يجب أن تُرجع السلسلة النصية `Not Equal`

```js
assert(compareEquality('20', 20) === 'Not Equal');
```

يجب أن تستخدم معامل `===`

```js
assert(__helpers.removeJSComments(code).match(/===/g));
```

# --seed--

## --seed-contents--

```js
// Setup
function compareEquality(a, b) {
  if (a == b) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

compareEquality(10, "10");
```

# --solutions--

```js
function compareEquality(a,b) {
  if (a === b) {
    return "Equal";
  }
  return "Not Equal";
}
```
