---
id: 5a2efd662fb457916e1fe604
title: التكرار باستخدام حلقات JavaScript من نوع Do...While
challengeType: 1
forumTopicId: 301172
dashedName: iterate-with-javascript-do---while-loops
---

# --description--

نوع الحلقة التالي الذي ستتعلمه يُسمى حلقة `do...while`. تُسمى حلقة `do...while` لأنها تبدأ أولًا بـ `do` تنفيذ مرة واحدة للكود داخل الحلقة مهما كان، ثم تستمر في تشغيل الحلقة `while` تحقق الشرط المحدد قيمة `true`.

```js
const ourArray = [];
let i = 0;

do {
  ourArray.push(i);
  i++;
} while (i < 5);
```

المثال أعلاه يتصرف بشكل مشابه لأنواع الحلقات الأخرى، وستبدو المصفوفة الناتجة مثل `[0, 1, 2, 3, 4]`. لكن ما يميز حلقة `do...while` عن الحلقات الأخرى هو كيف تتصرف عندما يفشل الشرط في أول تحقق. لنرَ هذا عمليًا. هنا حلقة `while` عادية ستشغل الكود داخل الحلقة طالما أن `i < 5`:

```js
const ourArray = []; 
let i = 5;

while (i < 5) {
  ourArray.push(i);
  i++;
}
```

في هذا المثال، نُهيئ قيمة `ourArray` إلى مصفوفة فارغة وقيمة `i` إلى 5. عند تنفيذ حلقة `while`، يُقيّم الشرط إلى `false` لأن `i` ليست أقل من 5، لذلك لا ننفذ الكود داخل الحلقة. النتيجة أن `ourArray` ستبقى بدون قيم مضافة وستظل تبدو مثل `[]` عندما ينتهي تنفيذ كل الكود في المثال أعلاه. الآن، ألقِ نظرة على حلقة `do...while`:

```js
const ourArray = []; 
let i = 5;

do {
  ourArray.push(i);
  i++;
} while (i < 5);
```

في هذه الحالة، نُهيئ قيمة `i` إلى 5، تمامًا كما فعلنا مع حلقة `while`. عند الوصول إلى السطر التالي، لا يوجد شرط للتقييم، لذلك ننتقل إلى الكود داخل الأقواس المتعرجة وننفذه. سنضيف عنصرًا واحدًا إلى المصفوفة ثم نزيد `i` قبل أن نصل إلى تحقق الشرط. عندما نُقيّم الشرط `i < 5` في السطر الأخير، نرى أن `i` أصبحت 6، وهو ما يفشل تحقق الشرط، فنخرج من الحلقة وننتهي. في نهاية المثال أعلاه، قيمة `ourArray` هي `[5]`. جوهريًا، تضمن حلقة `do...while` أن الكود داخل الحلقة سينفذ مرة واحدة على الأقل. دعنا نجرب تشغيل حلقة `do...while` بدفع قيم إلى مصفوفة.

# --instructions--

غيّر حلقة `while` في الكود إلى حلقة `do...while` بحيث تدفع الحلقة الرقم `10` فقط إلى `myArray`، ويكون `i` مساويًا لـ `11` عندما ينتهي تنفيذ كودك.

# --hints--

يجب أن تستخدم حلقة `do...while` لهذا التمرين.

```js
assert(__helpers.removeJSComments(code).match(/do/g));
```

يجب أن تكون قيمة `myArray` مساوية لـ `[10]`.

```js
assert.deepEqual(myArray, [10]);
```

يجب أن تكون قيمة `i` مساوية لـ `11`.

```js
assert.equal(i, 11);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];
let i = 10;

// Only change code below this line
while (i < 5) {
  myArray.push(i);
  i++;
}
```

# --solutions--

```js
const myArray = [];
let i = 10;
do {
  myArray.push(i);
  i++;
} while (i < 5)
```
