---
id: 5cc0bd7a49b71cb96132e54c
title: استخدام العودية لإنشاء نطاق من الأعداد
challengeType: 1
forumTopicId: 301180
dashedName: use-recursion-to-create-a-range-of-numbers
---

# --description--

مواصلةً للتحدٍّ السابق، نمنحك فرصة أخرى لإنشاء دالة عودية لحل مشكلة.

# --instructions--

عرّفنا دالة باسم `rangeOfNumbers` بمعلمتين. يجب أن تُرجع الدالة مصفوفة من أعداد صحيحة تبدأ برقم يمثله المعلم `startNum` وتنتهي برقم يمثله المعلم `endNum`. سيكون الرقم الابتدائي دائمًا أقل من أو يساوي الرقم النهائي. يجب أن تستخدم دالتك العودية باستدعاء نفسها دون استخدام أي حلقات. كما يجب أن تعمل في الحالات التي يكون فيها `startNum` و `endNum` متساويين.

# --hints--

يجب أن تُرجع دالتك مصفوفة.

```js
assert(Array.isArray(rangeOfNumbers(5, 10)));
```

يجب ألا يستخدم شفرتك أي صيغة حلقات (`for` أو `while` أو دوال مرتبة مثل `forEach`، `map`، `filter`، أو `reduce`).

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

يجب أن تستخدم `rangeOfNumbers` العودية (استدعاء نفسها) لحل هذا التحدٍّ.

```js
assert(
  rangeOfNumbers.toString().match(/rangeOfNumbers\s*\(.+\)/)
);
```

يجب أن تُرجع `rangeOfNumbers(1, 5)` `[1, 2, 3, 4, 5]`.

```js
assert.deepStrictEqual(rangeOfNumbers(1, 5), [1, 2, 3, 4, 5]);
```

يجب أن تُرجع `rangeOfNumbers(6, 9)` `[6, 7, 8, 9]`.

```js
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

يجب أن تُرجع `rangeOfNumbers(4, 4)` `[4]`.

```js
assert.deepStrictEqual(rangeOfNumbers(4, 4), [4]);
```

يجب عدم استخدام متغيرات عامة لتخزين المصفوفة مؤقتًا.

```js
rangeOfNumbers(1, 3)
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

# --seed--

## --seed-contents--

```js
function rangeOfNumbers(startNum, endNum) {
  return [];
};
```

# --solutions--

```js
function rangeOfNumbers(startNum, endNum) {
  if (endNum - startNum === 0) {
    return [startNum];
  } else {
    const numbers = rangeOfNumbers(startNum, endNum - 1);
    numbers.push(endNum);
    return numbers;
  }
}
```
