---
id: 5cd9a70215d3c4e65518328f
title: استخدام العودية لإنشاء عد تنازلي
challengeType: 1
forumTopicId: 305925
dashedName: use-recursion-to-create-a-countdown
---

# --description--

في <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/replace-loops-using-recursion" target="_blank" rel="noopener noreferrer nofollow">تحدٍّ سابق</a>، تعلمت كيف تستخدم العودية لتحل محل حلقة `for`. الآن، دعنا نلقي نظرة على دالة أكثر تعقيدًا تُرجع مصفوفة من الأعداد الصحيحة المتتالية تبدأ من `1` حتى الرقم الممرر إلى الدالة.

كما ذُكر في التحدي السابق، سيكون هناك <dfn>حالة أساسية</dfn>. تخبر الحالة الأساسية الدالة العودية متى لا تحتاج إلى استدعاء نفسها بعد الآن. هي حالة بسيطة تكون فيها قيمة الإرجاع معروفة مسبقًا. سيكون هناك أيضًا <dfn>استدعاء عودي</dfn> ينفذ الدالة الأصلية بمعلمات مختلفة. إذا كُتبت الدالة بشكل صحيح، فسيتم الوصول في النهاية إلى الحالة الأساسية.

على سبيل المثال، لنفترض أنك تريد كتابة دالة عودية تُرجع مصفوفة تحتوي الأعداد من `1` حتى `n`. يجب أن تقبل هذه الدالة معلمة واحدة، `n`، تمثل الرقم النهائي. ثم يجب أن تستدعي نفسها بقيم أصغر تدريجيًا من `n` حتى تصل إلى `1`. يمكنك كتابة الدالة كما يلي:

```javascript
function countup(n) {
  if (n < 1) {
    return [];
  } else {
    const countArray = countup(n - 1);
    countArray.push(n);
    return countArray;
  }
}
console.log(countup(5));
```

سيتم عرض القيمة `[1, 2, 3, 4, 5]` في وحدة التحكم.

في البداية، يبدو هذا غير بديهي لأن قيمة `n` *تنخفض*، لكن القيم في المصفوفة النهائية *تزداد*. يحدث هذا لأن عملية الإضافة تتم في النهاية، بعد أن يعود الاستدعاء العودي. في اللحظة التي تُضاف فيها `n` إلى المصفوفة، يكون `countup(n - 1)` قد تم تقييمه وأرجع `[1, 2, ..., n - 1]` بالفعل.

# --instructions--

لقد عرفنا دالة باسم `countdown` بمعلمة واحدة (`n`). يجب أن تستخدم الدالة العودية لإرجاع مصفوفة تحتوي الأعداد الصحيحة من `n` حتى `1` بناءً على المعلمة `n`. إذا استُدعي الدالة برقم أقل من 1، يجب أن تُرجع مصفوفة فارغة. على سبيل المثال، استدعاء هذه الدالة مع `n = 5` يجب أن يُرجع المصفوفة `[5, 4, 3, 2, 1]`. يجب أن تستخدم دالتك العودية باستدعاء نفسها ولا تستخدم أي نوع من الحلقات.

# --hints--

`countdown(-1)` يجب أن تُرجع مصفوفة فارغة.

```js
assert.isEmpty(countdown(-1));
```

`countdown(10)` يجب أن تُرجع `[10, 9, 8, 7, 6, 5, 4, 3, 2, 1]`

```js
assert.deepStrictEqual(countdown(10), [10, 9, 8, 7, 6, 5, 4, 3, 2, 1]);
```

`countdown(5)` يجب أن تُرجع `[5, 4, 3, 2, 1]`

```js
assert.deepStrictEqual(countdown(5), [5, 4, 3, 2, 1]);
```

يجب ألا يعتمد شفرتك على أي نوع من الحلقات (`for`، `while` أو الدوال عالية المستوى مثل `forEach`، `map`، `filter`، و `reduce`).

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

يجب أن تستخدم العودية لحل هذه المشكلة.

```js
assert(
  countdown.toString().match(/countdown\s*\(.+\)/)
);
```

يجب ألا تُستخدم المتغيرات العامة لتخزين المصفوفة مؤقتًا.

```js
countdown(1)
assert.deepStrictEqual(countdown(5), [5, 4, 3, 2, 1]);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
function countdown(n){
  return;
}
// Only change code above this line
```

# --solutions--

```js
function countdown(n){
   return n < 1 ? [] : [n].concat(countdown(n - 1));
}
```
