---
id: 5cfa3679138e7d9595b9d9d4
title: استبدال الحلقات باستخدام العودية
challengeType: 1
videoUrl: >-
  https://www.freecodecamp.org/news/how-recursion-works-explained-with-flowcharts-and-a-video-de61f40cb7f9/
forumTopicId: 301175
dashedName: replace-loops-using-recursion
---

# --description--

العودية هي مفهوم أن الدالة يمكن التعبير عنها باستخدام نفسها. لفهم ذلك، ابدأ بالتفكير في المهمة التالية: اضرب أول `n` عناصر من مصفوفة لإنشاء حاصل ضرب تلك العناصر. باستخدام حلقة `for`، يمكنك فعل ذلك:

```js
  function multiply(arr, n) {
    let product = 1;
    for (let i = 0; i < n; i++) {
      product *= arr[i];
    }
    return product;
  }
```

ومع ذلك، لاحظ أن `multiply(arr, n) == multiply(arr, n - 1) * arr[n - 1]`. هذا يعني أنه يمكنك إعادة كتابة `multiply` باستخدام نفسها دون الحاجة إلى استخدام حلقة.

```js
  function multiply(arr, n) {
    if (n <= 0) {
      return 1;
    } else {
      return multiply(arr, n - 1) * arr[n - 1];
    }
  }
```

النسخة العودية من `multiply` تتفكك هكذا. في <dfn>الحالة الأساسية</dfn>، حيث `n <= 0`، تُرجع 1. للقيم الأكبر من `n`، تستدعي نفسها ولكن مع `n - 1`. يتم تقييم استدعاء الدالة هذا بنفس الطريقة، مستدعية `multiply` مرة أخرى حتى يصبح `n <= 0`. عند هذه النقطة، يمكن لكل الدوال أن تُرجع والقيمة الأصلية لـ `multiply` تُرجع النتيجة.

**ملاحظة:** يجب أن تحتوي الدوال العودية على حالة أساسية تُرجع فيها دون استدعاء الدالة مرة أخرى (في هذا المثال، عندما `n <= 0`)، وإلا فلن تنتهي من التنفيذ أبدًا.

# --instructions--

اكتب دالة عودية، `sum(arr, n)`, تُرجع مجموع أول `n` عناصر من مصفوفة `arr`.

# --hints--

`sum([1], 0)` يجب أن تساوي 0.

```js
assert.equal(sum([1], 0), 0);
```

`sum([2, 3, 4], 1)` يجب أن تساوي 2.

```js
assert.equal(sum([2, 3, 4], 1), 2);
```

`sum([2, 3, 4, 5], 3)` يجب أن تساوي 9.

```js
assert.equal(sum([2, 3, 4, 5], 3), 9);
```

يجب ألا تعتمد شفرتك على أي نوع من الحلقات (`for` أو `while` أو الدوال عالية المستوى مثل `forEach`، `map`، `filter`، أو `reduce`).

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

يجب أن تستخدم العودية لحل هذه المشكلة.

```js
assert(
  sum.toString().match(/sum\(.*\)/g).length > 1
);
```

# --seed--

## --seed-contents--

```js
function sum(arr, n) {
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function sum(arr, n) {
  // Only change code below this line
  if(n <= 0) {
    return 0;
  } else {
    return sum(arr, n - 1) + arr[n - 1];
  }
  // Only change code above this line
}
```
