---
id: 5ee127a03c3b35dd45426493
title: تعيين قيمة متغير إلى متغير آخر
challengeType: 1
forumTopicId: 418265
dashedName: assigning-the-value-of-one-variable-to-another
---

# --description--

بعد تعيين قيمة إلى متغير باستخدام عامل <dfn>التعيين</dfn>، يمكنك تعيين قيمة ذلك المتغير إلى متغير آخر باستخدام عامل <dfn>التعيين</dfn>.

```js
var myVar;
myVar = 5;
var myNum;
myNum = myVar;
```

الشيفرة أعلاه تعلن عن متغير `myVar` بدون قيمة، ثم تُعيّن له القيمة `5`. بعد ذلك، يُعلن عن متغير باسم `myNum` بدون قيمة. ثم تُعيّن محتويات `myVar` (وهي `5`) إلى المتغير `myNum`. الآن، يحتوي `myNum` أيضًا على القيمة `5`.

# --instructions--

عيّن محتويات `a` إلى المتغير `b`.

# --hints--

لا يجب تغيير الشفرة فوق التعليق المحدد.

```js
assert(/var a;/.test(__helpers.removeJSComments(code)) && /a = 7;/.test(__helpers.removeJSComments(code)) && /var b;/.test(__helpers.removeJSComments(code)));
```

يجب أن تكون قيمة `b` هي `7`.

```js
assert(typeof b === 'number' && b === 7);
```

يجب تعيين `a` إلى `b` باستخدام `=`.

```js
assert(/b\s*=\s*a\s*/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
var a;
a = 7;
var b;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
var b;
b = a;
```
