---
id: bd7123c9c441eddfaeb4bdef
title: علّق شفرتك في JavaScript
challengeType: 1
forumTopicId: 16783
dashedName: comment-your-javascript-code
---

# --description--

التعليقات هي أسطر شفرة يتجاهلها JavaScript عمدًا. التعليقات طريقة رائعة لترك ملاحظات لنفسك وللآخرين الذين سيحتاجون لاحقًا لفهم ما تفعله تلك الشفرة.

هناك طريقتان لكتابة التعليقات في JavaScript:

استخدام `//` يخبر JavaScript بتجاهل الباقي من النص في السطر الحالي. هذا تعليق داخل السطر:

```js
// This is an in-line comment.
```

يمكنك إنشاء تعليق متعدد الأسطر يبدأ بـ `/*` وينتهي بـ `*/`. هذا تعليق متعدد الأسطر:

```js
/* This is a
multi-line comment */
```

**ملاحظة:** أثناء كتابة الشفرة، يجب أن تضيف تعليقات بانتظام لتوضيح وظيفة أجزاء الشفرة. التعليق الجيد يساعد في توصيل هدف الشفرة—لغيرك *ولنفسك* في المستقبل.

# --instructions--

جرّب إنشاء نوع واحد من كل نوع من التعليقات.

# --hints--

يجب أن تنشئ تعليقًا بأسلوب `//` يحتوي على خمسة أحرف على الأقل.

```js
assert(code.match(/(\/\/)...../g));
```

يجب أن تنشئ تعليقًا بأسلوب `/* */` يحتوي على خمسة أحرف على الأقل.

```js
assert(code.match(/(\/\*)([^\/]{5,})(?=\*\/)/gm));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
// Fake Comment
/* Another Comment */
```
