---
id: bd7993c9c69feddfaeb7bdef
title: الضرب بين عددين عشريين باستخدام JavaScript
challengeType: 1
forumTopicId: 301173
dashedName: multiply-two-decimals-with-javascript
---

# --description--

يمكنك في JavaScript إجراء عمليات حسابية على الأعداد العشرية تمامًا مثل الأعداد الصحيحة.

لنضرب عددين عشريين معًا للحصول على حاصل ضربهما.

# --instructions--

غيّر القيمة `0.0` بحيث يصبح حاصل الضرب مساويًا لـ `5.0`.

# --hints--

يجب أن تكون قيمة المتغير `product` مساوية لـ `5.0`.

```js
assert(product === 5.0);
```

يجب أن تستخدم العامل `*`

```js
assert(/\*/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const product = 2.0 * 0.0;
```

# --solutions--

```js
const product = 2.0 * 2.5;
```
