---
id: cf1111c1c11feddfaeb3bdef
title: جمع رقمين باستخدام JavaScript
challengeType: 1
forumTopicId: 16650
dashedName: add-two-numbers-with-javascript
---

# --description--

`Number` هو نوع بيانات في JavaScript يُستخدم لتمثيل البيانات الرقمية.

الآن جرّب جمع رقمين باستخدام JavaScript.

تستخدم JavaScript الرمز `+` كعامل جمع عندما يُوضع بين رقمين.

**مثال:**

```js
const myVar = 5 + 10;
```

الآن يحتوي `myVar` على القيمة `15`.

# --instructions--

غيّر القيمة `0` بحيث يكون مجموعها مساويًا لـ `20`.

# --hints--

يجب أن يكون `sum` مساويًا لـ `20`.

```js
assert(sum === 20);
```

يجب استخدام العامل `+`.

```js
assert(/\+/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const sum = 10 + 0;
```

# --solutions--

```js
const sum = 10 + 10;
```
