---
id: cf1111c1c11feddfaeb5bdef
title: التكرار باستخدام حلقات For في JavaScript
challengeType: 1
forumTopicId: 18219
dashedName: iterate-with-javascript-for-loops
---

# --description--

يمكنك تشغيل نفس الشفرة عدة مرات باستخدام حلقة.

أكثر أنواع الحلقات شيوعًا في JavaScript تُسمى حلقة `for` لأنها تعمل لعدد محدد من المرات.

تُعلن حلقات for بثلاثة تعبيرات اختيارية مفصولة بفواصل منقوطة:

`for (a; b; c)`، حيث `a` هو بيان التهيئة، و`b` هو بيان الشرط، و`c` هو التعبير النهائي.

يُنفذ بيان التهيئة مرة واحدة فقط قبل بدء الحلقة. يُستخدم عادة لتعريف وإعداد متغير الحلقة.

يُقيّم بيان الشرط في بداية كل تكرار للحلقة ويستمر طالما كانت نتيجته `true`. عندما يكون الشرط `false` في بداية التكرار، تتوقف الحلقة عن التنفيذ. هذا يعني أنه إذا بدأ الشرط كـ false، فلن تنفذ الحلقة أبدًا.

يُنفذ التعبير النهائي في نهاية كل تكرار للحلقة، قبل التحقق من الشرط التالي، وعادةً ما يُستخدم لزيادة أو نقصان عداد الحلقة.

في المثال التالي نُهيئ بـ `i = 0` ونتكرر طالما أن شرطنا `i < 5` صحيح. سنزيد `i` بمقدار `1` في كل تكرار للحلقة باستخدام `i++` كتعبير نهائي.

```js
const ourArray = [];

for (let i = 0; i < 5; i++) {
  ourArray.push(i);
}
```

سيكون لـ `ourArray` الآن القيمة `[0, 1, 2, 3, 4]`.

# --instructions--

استخدم حلقة `for` لإضافة القيم من 1 إلى 5 إلى `myArray`.

# --hints--

يجب أن تستخدم حلقة `for` لهذا الغرض.

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

يجب أن تساوي `myArray` القيمة `[1, 2, 3, 4, 5]`.

```js
assert.deepEqual(myArray, [1, 2, 3, 4, 5]);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 1; i < 6; i++) {
  myArray.push(i);
}
```
