---
id: 6626b4c58c027d86478ff5eb
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Die Lehrkraft ist wirklich zufrieden mit dem Programm, das Sie bisher erstellt haben. Jetzt möchte sie jedoch eine einfache Möglichkeit, um zu überprüfen, ob ein Schüler eine ausreichende Note hat. Eine ausreichende Note ist alles, was kein `"F"` ist.

Vervollständigen Sie die Funktion `hasPassingGrade`, die eine Schülerpunktzahl als Parameter erhält. Ihre Funktion soll `true` zurückgeben, wenn der Schüler eine ausreichende Note hat, und `false`, wenn nicht.

**Tipps**

- Verwenden Sie die `getGrade`-Funktion, um die Note des Schülers zu erhalten. Prüfen Sie dann, ob die Note ausreichend ist oder nicht.

# --hints--

Ihre `hasPassingGrade`-Funktion sollte einen Boolean-Wert zurückgeben.

```js
assert.strictEqual(typeof hasPassingGrade(100), "boolean");
```

Ihre `hasPassingGrade`-Funktion sollte `true` zurückgeben, wenn die Note ein `"A"` ist.

```js
assert.isTrue(hasPassingGrade(100));
```

Ihre `hasPassingGrade`-Funktion sollte `false` zurückgeben, wenn die Note ein `"F"` ist.

```js
assert.isFalse(hasPassingGrade(53));
```

Ihre `hasPassingGrade`-Funktion sollte `true` für alle ausreichenden Noten zurückgeben.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

--fcc-editable-region--
function hasPassingGrade(score) {
  
}


console.log(hasPassingGrade(100));
console.log(hasPassingGrade(53));
console.log(hasPassingGrade(87));
--fcc-editable-region--
```
