---
id: 6626b8dcf5057f896f948440
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Da die Lehrkraft nun alle benötigten Informationen hat, möchte sie dem Lernenden die Ergebnisse mitteilen können.

Vervollständigen Sie die `studentMsg`-Funktion mit `totalScores` und `studentScore` als Parameter. Die Funktion soll einen String zurückgeben, der eine Nachricht an den Lernenden darstellt.

Wenn der Lernende den Kurs bestanden hat, sollte der String folgendem Format entsprechen:

```md
Class average: average-goes-here. Your grade: grade-goes-here. You passed the course.
```

Wenn der Lernende den Kurs nicht bestanden hat, sollte der String folgendem Format entsprechen:

```md
Class average: average-goes-here. Your grade: grade-goes-here. You failed the course.
```

Ersetzen Sie `average-goes-here` durch den Durchschnitt der Gesamtnoten.  
Ersetzen Sie `grade-goes-here` durch die Note des Lernenden.

**Tipps**

- Verwenden Sie die `getAverage`-Funktion, um den Klassendurchschnitt zu ermitteln.  
- Verwenden Sie die `getGrade`-Funktion, um die Note des Lernenden zu ermitteln.  
- Verwenden Sie die String-Verkettung (`+`), um die Nachricht zu erstellen.  
- Achten Sie auf die Zeichensetzung und Leerzeichen in der Nachricht.

# --hints--

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` sollte die folgende Nachricht zurückgeben: `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` sollte die folgende Nachricht zurückgeben: `"Class average: 50.8. Your grade: A++. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A++. You passed the course.");
```

Ihre `studentMsg`-Funktion sollte basierend auf der Punktzahl des Lernenden und dem Klassendurchschnitt die korrekte Nachricht zurückgeben.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

--fcc-editable-region--
function studentMsg(totalScores, studentScore) {

}
console.log(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37));
--fcc-editable-region--

```


# --solutions--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

function studentMsg(totalScores, studentScore) {
  let average = getAverage(totalScores);
  let grade = getGrade(studentScore);

  return `Class average: ${average}. Your grade: ${grade}. You ${
    hasPassingGrade(studentScore) ? "passed" : "failed"
  } the course.`;
}

```
