---
id: 6626a060c4006f793e10cb33
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Jetzt braucht die Lehrkraft Ihre Hilfe, um die Punktzahl der Lernenden in eine Note umzuwandeln.

Vervollständigen Sie die `getGrade`-Funktion, die eine Zahl `score` als Parameter erhält. Ihre Funktion soll einen String zurückgeben, der eine Note basierend auf der Punktzahl darstellt.

Hier sind die Punktzahlen und ihre entsprechenden Noten:

| Punktzahlbereich | Note     |
| ----------- | ------- |
| `100`       | `"A++"` |
| `90 - 99`   | `"A"`   |
| `80 - 89`   | `"B"`   |
| `70 - 79`   | `"C"`   |
| `60 - 69`   | `"D"`   |
| `0 - 59`    | `"F"`   |


**Hinweise**

- Denken Sie daran, dass Sie bereits über bedingte Anweisungen (`if`, `else if` und `else`) gelernt haben.
- Denken Sie daran, dass Sie bereits Vergleichsoperatoren (`>`, `<`, `>=`, `<=`, `===`) kennengelernt haben.
- Denken Sie daran, dass Sie bereits den logischen UND-Operator (`&&`) kennengelernt haben.

# --hints--

Ihre `getGrade`-Funktion sollte `"A++"` zurückgeben, wenn die Punktzahl `100` ist.

```js
assert.strictEqual(getGrade(100), "A++");
```

Ihre `getGrade`-Funktion sollte `"A"` zurückgeben, wenn die Punktzahl `94` ist.

```js
assert.strictEqual(getGrade(94), "A");
```

Ihre `getGrade`-Funktion sollte `"B"` zurückgeben, wenn die Punktzahl zwischen `80` und `89` liegt.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

Ihre `getGrade`-Funktion sollte `"C"` zurückgeben, wenn die Punktzahl `78` ist.

```js
assert.strictEqual(getGrade(75), "C");
```

Ihre `getGrade`-Funktion sollte `"D"` zurückgeben, wenn die Punktzahl zwischen `60` und `69` liegt.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

Ihre `getGrade`-Funktion sollte `"F"` zurückgeben, wenn die Punktzahl `57` ist.

```js
assert.strictEqual(getGrade(57), "F");
```

Ihre `getGrade`-Funktion sollte `"F"` zurückgeben, wenn die Punktzahl zwischen `0` und `59` liegt.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}
--fcc-editable-region--
function getGrade(score) {

}

console.log(getGrade(96));
console.log(getGrade(82));
console.log(getGrade(56));
--fcc-editable-region--
```
