---
id: 662693f82c91a66be46c881b
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

Eine Lehrkraft hat die Tests ihrer Schüler bewertet und benötigt Ihre Hilfe, um die Durchschnittsnote der Klasse zu berechnen.

Vervollständigen Sie die `getAverage`-Funktion, die ein Array von Testergebnissen entgegennimmt und die Durchschnittsnote zurückgibt.

Der Durchschnitt wird berechnet, indem alle Noten addiert und durch die Gesamtanzahl der Noten geteilt werden.

```md
average = sum of all scores / total number of scores
```

Einige Funktionsaufrufe wurden für Sie bereitgestellt, damit Sie Ihren Code testen können.

**Tipps**

- Sie können eine Schleife verwenden, um das `scores`-Array zu durchlaufen und alle Noten zu addieren.
- Sie können die `length`-Eigenschaft verwenden, um die Gesamtanzahl der Noten zu erhalten.

# --hints--

Ihre `getAverage`-Funktion sollte eine Zahl zurückgeben.

```js
assert.strictEqual(typeof getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 'number');
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` sollte `71.7` zurückgeben.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` sollte `85.4` zurückgeben.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` sollte `92.4` zurückgeben.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

Ihre `getAverage`-Funktion sollte die Durchschnittsnote der Testergebnisse zurückgeben.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function getAverage(scores) {

}

console.log(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
console.log(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]));
--fcc-editable-region--
```
