---
id: 69899123889dcf5b19be1847
title: "implementare l'algoritmo di ricerca in profondità"
challengeType: 26
dashedName: implement-depth-first-search-algorithm
---

# --description--

In questo laboratorio implementerai un algoritmo di attraversamento di un grafo chiamato <dfn>ricerca in profondità</dfn> (DFS).

Mentre una <dfn>ricerca in ampiezza</dfn> (BFS) esplora i vicini livello per livello, la <dfn>ricerca in profondità</dfn> scende il più possibile lungo un singolo percorso di archi prima di tornare indietro.

Quando l'algoritmo raggiunge la fine di un percorso (un nodo senza vicini non visitati), torna indietro al nodo più recente che ha ancora archi inesplorati e continua la ricerca da lì.

Per implementare questo algoritmo in modo iterativo, vorrai usare una **stack**. In JavaScript, un array standard `Array` funziona perfettamente come stack usando i metodi `.push()` e `.pop()`, seguendo il principio **Last-In-First-Out (LIFO)**. Questo garantisce che il vicino aggiunto più di recente allo stack sia il prossimo da esplorare.

Un output semplice di questo algoritmo è una lista di nodi raggiungibili da un nodo dato. Perciò, dovrai anche tenere traccia dei nodi che visiti.

**Obiettivo**: Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti definire una funzione chiamata `dfs`.
1. La funzione `dfs` dovrebbe prendere due argomenti:

    - `graph`: Una matrice di adiacenza (un array di array).
    - `root`: Un'etichetta numerica del nodo (l'indice di partenza) che è il valore numerico del nodo tra `0` e `n - 1`, dove `n` è il numero totale di nodi nel grafo.

1. La funzione dovrebbe restituire un array contenente tutte le etichette dei nodi raggiungibili dal nodo di partenza.

# --hints--

Dovresti avere una funzione chiamata `dfs` che prende due argomenti.

```js
assert.isFunction(dfs);
assert.lengthOf(dfs, 2);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)` dovrebbe restituire una lista con `1`, `2`, `3` e `0`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)` dovrebbe restituire una lista con `1`, `2`, `3` e `0`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)` dovrebbe restituire `[3]`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3), [3]);
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)` dovrebbe restituire una lista con `3` e `2`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3), [3, 2]);

```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)` dovrebbe restituire una lista con `0` e `1`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0), [0, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfs(graph, root) {
  let stack = [root];
  let visited = [];

  while (stack.length > 0) {
    let current = stack.pop();

    if (!visited.includes(current)) {
      visited.push(current);
      let neighbors = graph[current];
      
      for (let i = 0; i < neighbors.length; i++) {
        if (neighbors[i] === 1) {
          stack.push(i);
        }
      }
    }
  }
  return visited;
}
```
