---
id: 6582267ef39f5047411c1743
title: Hatua 14
challengeType: 20
dashedName: step-14
---

# --description--

Ifafanue kitendakazi kinachoitwa `print_expenses` kinachopokea kigezo kimoja `expenses`. Kitendakazi hiki baadaye kitatumika kuonyesha kila gharama katika orodha yako.

Jaza mwili wa kitendakazi chako kipya na usemi wa `pass`.

# --hints--

Unapaswa kufafanua kitendakazi kipya kinachoitwa `print_expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(print_expenses)
  `))
})
```

Kitendakazi chako cha `print_expenses` kinapaswa kupokea kigezo kimoja `expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(print_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
