---
id: 658212ba98182f3e855e85f9
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

Katika mradi huu, utajenga kifuatiliaji cha matumizi rahisi, lakini chenye utendaji, kwa kutumia Python.

Anza kwa kutangaza kitendakazi kinachoitwa `add_expense` kinachopokea vigezo vitatu: `expenses`, `amount` na `category`. Tumia neno kuu la `pass` kujaza mwili wa kitendakazi.

# --hints--

Unapaswa kutumia neno kuu la `def` kutangaza kitendakazi kinachoitwa `add_expense`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(add_expense)
  `))
})
```

Kitendakazi chako cha `add_expense` kinapaswa kupokea vigezo vitatu: `expenses`, `amount` na `category`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(add_expense))
    sig == '(expenses, amount, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
