---
id: 65821fcc010c3245718f2a06
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Kigezo cha `expenses` cha kitendakazi chako cha `add_expense` kitakuwa orodha ya matumizi. Unataka uweze kuongeza vitu mwishoni mwa orodha yako. Kwa hiyo utatumia njia ya orodha ya `.append()`.

Ongeza wito wa njia ya `.append()` kwenye orodha ya `expenses`. Usipitisha hoja yoyote kwa `.append()` kwa sasa.

# --hints--

Unapaswa kuwa na `expenses.append()` katika kitendakazi chako.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append()")`)) })
```

Hupaswi kuwa na `pass` katika kitendakazi chako.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("add_expense").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    pass
--fcc-editable-region--

expenses = []
```
