---
id: 65823040bdb32949d4c1a96f
title: Hatua 18
challengeType: 20
dashedName: step-18
---

# --description--

Utahitaji kitendakazi kuhesabu jumla ya gharama.

Fafanua kitendakazi kinachoitwa `total_expenses` kinachopokea kigezo kimoja `expenses`. Jaza mwili wa kitendakazi kwa usemi wa `pass` kwa sasa.

# --hints--

Unapaswa kufafanua kitendakazi kinachoitwa `total_expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(total_expenses)
  `))
})
```

Kitendakazi chako cha `total_expenses` kinapaswa kupokea kigezo kimoja cha `expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(total_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
