---
id: 65824561f3478e5371a33ae5
title: Hatua 32
challengeType: 20
dashedName: step-32
---

# --description--

Hatua inayofuata ni kufafanua kitendakazi kikuu, ambacho kitakuwa sehemu ya kuingia ya programu ya kufuatilia matumizi ya kiingilio.

Fafanua kitendakazi chenye jina `main` bila vigezo. Jaza mwili wa kitendakazi hicho na orodha ya `expenses` uliyounda mwanzoni mwa mradi huu. Utatumia orodha hii kuhifadhi rekodi za matumizi.

# --hints--

Unapaswa kufafanua kitendakazi chenye jina `main()` bila vigezo.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(main)
    sig = str(inspect.signature(main))
    sig == '()'
  `))
})
```

Unapaswa kuhamisha orodha ya `expenses` ndani ya mwili wa kitendakazi cha `main()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_body().is_equivalent("expenses = []")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--

expenses = []
--fcc-editable-region--
```
