---
id: 65822308ded1ba4632f66aa3
title: Hatua 12
challengeType: 20
dashedName: step-12
---

# --description--

Kamusi ni aina nyingine ya data iliyojengwa ndani katika Python. Kamusi ni mkusanyiko wa data katika muundo wa jozi za ufunguo-thamani. Kamusi huundwa kwa mabano ya mviringo (`{}`) na zina jozi za ufunguo-thamani zilizotenganishwa kwa koma. Kila ufunguo hufuatiwa na nukta (`:`) na thamani:

```py
{'amount': 50.0, 'category': 'Food'}
```

Katika mfano uliotangulia, `'amount'` na `'category'` ni funguo, na `50.0` na `'Food'` ni thamani zao zinazolingana.

Unda kamusi yenye ufunguo `'amount'` na thamani ya kigezo `amount` kisha pita kamusi yako mpya kwa wito wa `.append()`.

# --hints--

Unapaswa kupitisha kamusi kama hoja kwa njia ya `.append()`.

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict)
`)) })
```

Unapaswa kupitisha `{'amount': amount}` kama hoja kwa njia ya `.append()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append()
--fcc-editable-region--

expenses = []
```
