---
id: 65823634a308cb4d6bcc5fa6
title: Hatua 21
challengeType: 20
dashedName: step-21
---

# --description--

Vitendakazi vya lambda vinaweza kuunganishwa vyema na kitendakazi cha `map()`, ambacho hutekeleza kitendakazi kilichobainishwa kwa kila kipengele katika mkusanyiko wa vitu, kama vile orodha:

```py
map(lambda x: x * 2, [1, 2, 3])
```

Kitendakazi kinachotekelezwa hupitishwa kama hoja ya kwanza, na inayoweza kurudiwa hupitishwa kama hoja ya pili.

Matokeo ya mfano hapo juu yatakuwa `[2, 4, 6]`, ambapo kila kipengele katika orodha iliyopitishwa kwa `map()` kimezidishwa mara mbili kwa kitendo cha kitendakazi cha lambda.

Badilisha wito wako wa `print()` ili kuchapisha matokeo ya wito wa `map()` ukiwa na `test` kama hoja ya kwanza, na `[2, 3, 5, 8]` kama hoja ya pili. Hutaweza kuona matokeo yanayoweza kusomeka bado.

# --hints--

Unapaswa kuwa na `print(map(test, [2, 3, 5, 8]))` katika msimbo wako.

```js
({ test: () => assert.match(code, /^print\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(test(3))
--fcc-editable-region--

expenses = []
```
