---
id: 65824dfdb6815d563b2d3256
title: Hatua 37
challengeType: 20
dashedName: step-37
---

# --description--

Kitendakazi cha `input()` kinachukua ingizo la mtumizi na hurudisha ingizo hilo katika mfumo wa mfuatano wa herufi.

Ndani ya mzunguko wako wa `while`, fanya wito wa kitendakazi cha `input()` ukipitisha mfuatano wa herufi `'Enter your choice: '` kama hoja, na weka thamani ya matokeo katika kigezo kinachoitwa `choice`.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `choice` ndani ya mzunguko wako wa `while`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].has_variable("choice")
`)) })
```

Unapaswa kuweka thamani ya matokeo ya kitendakazi cha `input()`, ukiwa na mfuatano wa herufi `'Enter your choice: '` kama hoja yake, katika kigezo chako cha `choice`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_variable("choice").is_equivalent("choice = input('Enter your choice: ')")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
--fcc-editable-region--
```
