---
id: 6582507654b3ed5712341382
title: Hatua 38
challengeType: 20
dashedName: step-38
---

# --description--

Utakuwa unatumia sentensi za masharti kukagua chaguo la mtumizi. Ikiwa chaguo ni `'1'`, inamaanisha mtumizi anataka kuongeza gharama.

Bado ukiwa ndani ya mzunguko wa `while`, chini ya kigezo cha `choice`, andika sentensi ya `if` ili kukagua kama `choice` ni sawa na mfuatano wa herufi `'1'`. Ikiwa ni kweli, itakuwa mwanzo wa kuongeza gharama mpya.

Ndani ya mwili wa sentensi ya `if`, tangaza kigezo cha `amount` na weka thamani ya wito tupu wa `input()`.

# --hints--

Unapaswa kuunda sentensi ya `if` inayokagua kama `choice` ni sawa na `'1'` katika mzunguko wako wa `while`.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
cond.is_equivalent("choice == '1'") or cond.is_equivalent("'1' == choice")
`)) })
```

Unapaswa kutangaza kigezo chenye jina `amount` ndani ya mwili wa sentensi yako ya `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("amount")
`)) })
```

Unapaswa kuweka thamani ya wito tupu wa `input()` kwa kigezo chako cha `amount`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("amount").is_equivalent("amount = input()")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')

--fcc-editable-region--
```
