---
id: 65821c1c40ccac44adf38e6b
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Tengeneza orodha tupu iitwayo `expenses`. Utaitumia kuhifadhi kila mojawapo ya matumizi yako.

# --hints--

Unapaswa kutangaza kigezo kipya kiitwacho `expenses`.

```js
({ test: () => assert(__userGlobals.has("expenses")) })
```

Unapaswa kuweka thamani ya orodha tupu katika kigezo `expenses`.

```js
({ test: () => assert.equal(__userGlobals.get("expenses").length, 0) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    pass
    
--fcc-editable-region--

--fcc-editable-region--
```
