---
id: 6582352cbdbcde4d0c4f7b0b
title: Hatua 20
challengeType: 20
dashedName: step-20
---

# --description--

Ili kuita kitendakazi cha lambda unaweza kutumia sintaksia ya kawaida ya kitendakazi na mabano ya kawaida baada ya jina la kigezo.

Ita kitendakazi chako cha `test` lambda na pita `3` kama hoja. Kisha, chapisha matokeo.

# --hints--

Unapaswa kuwa na `print(test(3))` katika msimbo wako.

```js
({ test: () => assert.match(code, /^print\s*\(\s*test\s*\(\s*3\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2

--fcc-editable-region--

expenses = []
```
