---
id: 65825b758fe85a5cebabc8c5
title: Hatua 45
challengeType: 20
dashedName: step-45
---

# --description--

Ili kuonyesha jumla ya matumizi, tengeneza kifungu cha `elif` kinachokagua kama `choice == '3'`.

Ikiwa ni kweli, ina maana mtumizi anataka kuona jumla ya matumizi. Hivyo fanya wito wa `print()` na pita mfuatano wa herufi `'\nTotal Expenses: '` kama hoja ya kwanza na `total_expenses(expenses)` kama hoja ya pili.

# --hints--

Unapaswa kuunda kifungu cha `elif` kinachokagua kama `choice == '3'`.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[2]
cond.is_equivalent("choice == '3'") or cond.is_equivalent("'3' == choice")
`)) })
```

Unapaswa kuchapisha `'\nTotal Expenses: ', total_expenses(expenses)` katika kifungu chako kipya cha `elif`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[2].has_call("print('\\\\nTotal Expenses: ', total_expenses(expenses))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
       
        choice = input('Enter your choice: ')

        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            add_expense(expenses, amount, category)
--fcc-editable-region--
        elif choice == '2':
            print('\nAll Expenses:')
            print_expenses(expenses)
        
--fcc-editable-region--
```
