---
id: 65824111a09164518320088d
title: Hatua 30
challengeType: 20
dashedName: step-30
---

# --description--

Kitendakazi cha `filter()` kinakuwezesha kuchagua vitu kutoka kwenye inayoweza kurudiwa, kama vile orodha, kulingana na matokeo ya kitendakazi:

```py
filter(my_function, my_list)
```

`filter()` huchukua kitendakazi kama hoja yake ya kwanza na inayoweza kurudiwa kama hoja yake ya pili. Hurejesha kielekezi, ambacho ni kitu maalum kinachokuwezesha kurudia vitu vya mkusanyiko, kama orodha.

Matokeo ya mfano hapo juu ni kielekezi chenye vitu vya `my_list` ambavyo `my_function` hurejeshea `True`.

Ndani ya kitendakazi cha `filter_expenses_by_category`, fanya wito wa `filter()` ukipitisha kitendakazi cha `lambda` ulichokiandika katika hatua iliyopita kama hoja ya kwanza na orodha ya `expenses` kama hoja ya pili.

# --hints--

Unapaswa kufanya wito wa `filter()` ndani ya kitendakazi cha `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")) == 1`)) })
```

Unapaswa kupitisha `lambda expense: expense['category'] == category` kama hoja ya kwanza kwa wito wa `filter()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[0].is_equivalent("lambda expense: expense['category'] == category")`)) })
```

Unapaswa kupitisha `expenses` kama hoja ya pili kwa wito wa `filter()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    lambda expense: expense['category'] == category
--fcc-editable-region--

expenses = []
```
