---
id: 6582575b8089f85b8b92d7c8
title: Hatua 42
challengeType: 20
dashedName: step-42
---

# --description--

Mara baada ya kupata maelezo ya gharama, unahitaji kuita kitendakazi cha `add_expense` kuongeza gharama mpya kwenye orodha ya `expenses`.

Baada ya kupata `amount` na `category` kwa kutumia `input()`, ita wito wa kitendakazi cha `add_expense`, ukipitisha hoja tatu: `expenses`, `amount` na `category`. 

- `expenses` ni orodha tupu iliyoundwa katika kitendakazi kikuu awali katika mradi huu.
- `amount` ni kiasi cha gharama.
- `category` ni kategoria ya gharama.

# --hints--

Unapaswa kuita kitendakazi cha `add_expense()` ndani ya tamko lako la `if`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_calls("add_expense")) == 1`)) })
```

Unapaswa kupitisha hoja tatu kwa wito lako la `add_expense()`: `expenses`, `amount`, na `category`. Mpangilio ni muhimu.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("add_expense(expenses, amount, category)")`)) })
```

Wito lako la `add_expense()` linapaswa kuja baada ya mgawo wa `amount` na `category`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].is_equivalent("amount = float(input('Enter amount: '))\\ncategory = input('Enter category: ')\\nadd_expense(expenses, amount, category)")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            
--fcc-editable-region--
```
