---
id: 65822bd82d708c4895080c35
title: Hatua 17
challengeType: 20
dashedName: step-17
---

# --description--

Katika Python, jambo muhimu kujua ni kwamba aina ile ile ya nukuu inayotumika kufafanua mfuatano wa herufi haiwezi kutumika ndani yake. Kwa mfano, mfuatano wa herufi `'I'm a string!'` si halali. Ili kutumia nukuu moja ndani ya mfuatano huo wa herufi unapaswa:

- Kuepuka nukuu hiyo kwa kuweka alama ya msalaba kabla yake: `'I\'m a string!'`
- Au tumia alama za nukuu maradufu kufafanua mfuatano wa herufi: `"I'm a string!"` (inapendekezwa).

Unaweza kupata thamani katika kamusi kupitia funguo zake. Unahitaji kutumia uandishi wa mabano na kuweka funguo kati ya mabano ya mraba:

```py
my_dict = {'amount': 50.0, 'category': 'Food'}
my_dict['amount'] # 50.0
```

Kwa sasa unafanya interpolation ya kamusi ya `expense` katika f-string yako. Badilisha usemi wa f-string ili upate thamani ya funguo ya `'amount'` na funguo ya `'category'` katika kamusi ya `expense`.

# --hints--

Unapaswa kupitisha `f'Amount: {expense["amount"]}, Category: {expense["category"]}'` kwa wito lako la `print()`. Kumbuka kutumia alama za nukuu maradufu ndani ya f-string yenye nukuu moja na kinyume chake.

```js
({ test: () => assert.match(code, /^\s+print\s*\(\s*f("|')Amount: \{\s*expense\s*\[\s*(?=[^\1])("|')amount\2\s*\]\s*\}, Category: \{\s*expense\s*\[\s*(?=[^\1])("|')category\3\s*\]\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense}, Category: {expense}')
--fcc-editable-region--

expenses = []
```
