---
id: 65823cfc74aa564ffc460489
title: Hatua 26
challengeType: 20
dashedName: step-26
---

# --description--

Sasa, fanya wito la `map()` ukitumia kitendakazi chako cha `lambda` kama hoja ya kwanza na orodha ya `expenses` kama hoja ya pili.

# --hints--

Unapaswa kufanya wito la kitendakazi cha `map()` ndani ya kitendakazi cha `total_expenses`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("total_expenses").find_calls("map")) == 1`)) })
```

Unapaswa kupitisha kitendakazi chako cha `lambda` kama hoja ya kwanza kwa wito la `map()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[0].is_equivalent("lambda expense: expense['amount']")`)) })
```

Unapaswa kupitisha orodha ya `expenses` kama hoja ya pili kwa wito la `map()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    lambda expense: expense['amount']
--fcc-editable-region--

expenses = []
```
