---
id: 65822773f4c60a479ec89f66
title: Hatua 15
challengeType: 20
dashedName: step-15
---

# --description--

Ndani ya kitendakazi cha `print_expenses`, tengeneza mzunguko wa `for` unaorudia kila kipengele katika orodha ya `expenses`. Tumia `expense` kama kigezo cha mzunguko na hamisha `pass` ndani ya mwili wa mzunguko.

# --hints--

Unapaswa kutengeneza mzunguko wa `for` unaorudia `expenses` katika kitendakazi chako cha `print_expenses`. Kumbuka kutumia `pass` ndani ya mwili wa mzunguko.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_iter().is_equivalent("expenses")`)) })
```

Unapaswa kutumia `expense` kama kigezo cha mzunguko.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_vars().is_equivalent("expense")`)) })
```

Unapaswa kuwa na `pass` ndani tu ya mwili wa mzunguko wako.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    pass
--fcc-editable-region--

expenses = []
```
