---
id: 65822934c8a0454803e3ef12
title: Hatua 16
challengeType: 20
dashedName: step-16
---

# --description--

Ifuatayo, utaonyesha maelezo ya kila matumizi.

Ndani ya mzunguko wa `for`, badilisha `pass` kwa wito wa `print()` na upitie f-string ifuatayo: `f'Amount: {expense}, Category: {expense}'`.

# --hints--

Unapaswa kuchapisha `f'Amount: {expense}, Category: {expense}'` katika mzunguko wako wa `for`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].has_call("print(f'Amount: {expense}, Category: {expense}')")`)) })
```

Haupaswi kuwa na `pass` ndani ya mwili wa mzunguko wako.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_bodies()[0].has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        pass
--fcc-editable-region--

expenses = []
```
