---
id: 6582324b5d37fe4b622bc3c0
title: Hatua 19
challengeType: 20
dashedName: step-19
---

# --description--

Kitendakazi cha lambda ni kitendakazi kifupi, kisicho na jina katika Python, kinachofaa kwa kazi rahisi za mara moja tu. Hutatangazwa kwa neno kuu la `lambda`, na hutumia sintaksia ifuatayo:

```py
lambda x: expr
```

Katika mfano ulio juu, `x` inawakilisha kigezo kinachotumika katika usemi wa kihesabu `expr`, na hufanya kazi kama kigezo chochote katika kitendakazi cha kawaida. `expr` ni usemi wa kihesabu unaotathminiwa na kurudishwa wakati kitendakazi cha lambda kinapoitwa.

Tengeneza kigezo kinachoitwa `test` na weka thamani yake kuwa kitendakazi cha lambda kinachopokea kigezo `x` na kurudisha `x * 2`.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `test`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("test")`))})
```

Unapaswa kuweka thamani ya `lambda x: x * 2` kwa kigezo chako cha `test`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("test").is_equivalent("test = lambda x: x * 2")`))})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
