---
id: 658246c28575d653d1f89b59
title: Hatua 33
challengeType: 20
dashedName: step-33
---

# --description--

Mzunguko wa `while` ni aina nyingine ya mzunguko unaotekeleza sehemu ya msimbo mradi tu hali fulani iwe `True`. Mzunguko unamalizika wakati hali inapo kuwa `False`:

```py
while condition:
    <code>
```

Chini ya orodha ya `expenses`, tengeneza mzunguko wa `while`. Tumia `True` kama hali, na chapisha mfuatano wa herufi `'\nExpense Tracker'` ndani ya mwili wa mzunguko kuonyesha kichwa cha programu.

# --hints--

Unapaswa kutengeneza mzunguko wa `while` ukitumia `True` kama hali ndani ya kitendakazi chako cha `main()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_conditions()[0].is_equivalent("True")`)) })
```

Unapaswa kuchapisha `'\nExpense Tracker'` katika mzunguko wako wa `while`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")[0].is_equivalent("print('\\\\nExpense Tracker')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--
def main():
    expenses = []
    
--fcc-editable-region--
```
