---
id: 65823ff0d4b991510fade1a8
title: Hatua 29
challengeType: 20
dashedName: step-29
---

# --description--

Ndani ya kitendakazi cha `filter_expenses_by_category`, badilisha `pass` na kitendakazi cha `lambda`. Tumia `expense` kama kigezo na tathmini kulinganisha kati ya thamani ya ufunguo wa `'category'` wa kamusi ya `expense` na `category` kwa kutumia opereta ya usawa.

# --hints--

Unapaswa kuunda kitendakazi cha `lambda` kinachotumia kigezo kinachoitwa `expense` na kinachothamini usemi wa kihesabu `expense['category'] == category` ndani ya kitendakazi cha `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_stmt("lambda expense: expense['category'] == category")`)) })
```

Haupasi kuwa na `pass` ndani ya kitendakazi cha `filter_expenses_by_category`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    pass
--fcc-editable-region--

expenses = []
```
