---
id: 65823e87c9741750a22085a7
title: Hatua 28
challengeType: 20
dashedName: step-28
---

# --description--

Ifuatayo, fafanua kitendakazi kinachoitwa `filter_expenses_by_category` kinachopokea vigezo viwili: `expenses` na `category`. Tumia `pass` kujaza mwili wa kitendakazi.

# --hints--

Unapaswa kufafanua kitendakazi kinachoitwa `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(filter_expenses_by_category)
  `))
})
```

Kitendakazi chako cha `filter_expenses_by_category` kinapaswa kupokea `expenses` na `category` kama vigezo.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(filter_expenses_by_category))
    sig == '(expenses, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
