---
id: 65823dde36392f505a39f7c7
title: Hatua 27
challengeType: 20
dashedName: step-27
---

# --description--

Mwishowe, pitisha wito wako wa `map()` kwa kitendakazi cha `sum()` kupata jumla ya matumizi na rudisha matokeo.

# --hints--

Unapaswa kuita kitendakazi cha `sum()` ukimpa wito wako wa sasa wa `map()` kama hoja na kurudisha matokeo kutoka kwa kitendakazi chako cha `total_expenses`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").has_return("sum(map(lambda expense: expense['amount'], expenses))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    map(lambda expense: expense['amount'], expenses)
--fcc-editable-region--

expenses = []
```
