---
id: 666951d2e4b27e3a3c72ca5f
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Unaweza pia kupata kipengele kimoja tu ili kupata thamani yake.

Orodha zina kielezo cha sifuri kama vile mfuatano wa herufi ulivyo. Hii ina maana kwamba kipengele cha kwanza kiko kwenye kielezo `0`, kipengele cha pili kiko kwenye kielezo `1` na kadhalika.

Ili kupata kipengele unatumia uandishi wa mabano. Kwa mfano, `example_list[1]` ni kupata kipengele kilicho kwenye kielezo `1`, kipengele cha pili, cha `example_list`.

Chapisha kipengele cha kwanza cha `my_list`.

# --hints--

Unapaswa kuchapisha kipengele cha kwanza cha orodha kwa kutumia `print(my_list[0])`.

```js
({
    test: () => runPython(`
calls =  _Node(_code).find_calls("print")
assert any(call.is_equivalent('print(my_list[0])') for call in calls), "print of my_list[0] not found"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)
--fcc-editable-region--
```
