---
id: 66694eaaeb7d253049fa7883
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Ili kuongeza vitu kwenye orodha, unaweza kuviandika kati ya mabano ya mraba, ukitenganisha thamani kwa koma.

Ongeza nambari `1` na `2` ndani ya orodha.

# --hints--

Kipengele cha kwanza cha `my_list` kinapaswa kuwa `1`.

```js
({
    test: () => runPython(`
        assert my_list[0] == 1, "my_list[0] is not 1"
    `)
})
```

Kipengele cha pili cha `my_list` kinapaswa kuwa `2`.

```js
({
    test: () => runPython(`
        assert my_list[1] == 2, "my_list[1] is not 2"
    `)
})
```

`my_list` inapaswa kuwa na vipengele viwili ndani yake.

```js
({
    test: () => runPython(`
        assert len(my_list) == 2, "my_list doesn't have two elements"
    `)
})
```

Unapaswa kuweka thamani ya `[1, 2]` kwa `my_list` moja kwa moja.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = []
--fcc-editable-region--
```
