---
id: 66694fc4bba24f33ca01fa5b
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Python ina njia chache za orodha. Kama njia za kuongeza au kuondoa vitu kwenye orodha.

Unaweza kuongeza kipengee mwishoni mwa orodha kwa kutumia njia ya `append()`. Kwa mfano:

```py
example_list = [4, 5, 6]

example_list.append(7)
```

`example_list` sasa ni `[4, 5, 6, 7]`.

Jaribu kutumia njia ya `append()` kuongeza `3` kwenye `my_list`. Kisha chapisha orodha.

# --hints--

Kipengee cha mwisho cha `my_list` kinapaswa kuwa `3`.

```js
({
    test: () => runPython(`
        assert my_list[2] == 3, "my_list[2] is not 3"
    `)
})
```

`my_list` inapaswa kuwa na urefu wa 3.

```js
({
    test: () => runPython(`
        assert len(my_list) == 3, "my_list doesn't have three elements"
    `)
})
```

Haupasi kubadilisha tamko la orodha.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

Unapaswa kuchapisha orodha.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls("print")
    assert any(call.is_equivalent('print(my_list)') for call in calls)
    `)
})
```

Unapaswa kuchapisha orodha baada ya kutumia `append()`.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.append(3)', 'print(my_list)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]
--fcc-editable-region--
```
