---
id: 66694bfa435cef241f6f9a68
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Kabla ya kuanza na mradi, utajifunza kuhusu orodha.

Tofauti na nambari na mfuatano wa herufi, <dfn>orodha</dfn> ni aina ya data inayofanya kazi kama kontena la thamani nyingine:

```py
example_list = [1, 2, 3, 4]
empty_list = []
```

Orodha ina sifa ya mabano ya mraba yanayozunguka thamani zote, na koma kati ya thamani, kama: `["A", "happy", "list"]`. Ikiwa orodha haina thamani yoyote, basi ni orodha tupu: `[]`.

Orodha inaweza kuwa na aina tofauti za data: `[1, "Up", ["Down", "Twice"]]`. Hii inajumuisha aina zote zinazowezekana za data. Pia inaweza kuwa na orodha nyingine ndani yake!

Tengeneza kigezo kinachoitwa `my_list` na umweke thamani ya orodha tupu.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `my_list`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('my_list'), "variable my_list not found"
    `)
})
```

Kigezo `my_list` kinapaswa kuwa na thamani ya orodha tupu.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent('my_list = []'), "variable my_list is not an empty list"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
