---
id: 666aae6a5d30a71f1fd7749f
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Njia ya `insert` inaweza kuongeza kipengele katika nafasi yoyote kwenye orodha. Hoja ya kwanza ni nafasi ambayo kipengele kinapaswa kuongezwa, na hoja ya pili ni kipengele cha kuongezwa. Kwa mfano, hivi ndivyo ya kuongeza kipengele kipya katika nafasi ya tatu ya `example_list`:

```py
example_list = [4, 5, 6, 7]

example_list.insert(2, 5.5)

print(example_list) # [4, 5, 5.5, 6, 7]
```

Ukitegemea `insert()`, ongeza `1` kwenye `my_list` katika nafasi sahihi ili kuhesabu kwa kuongezeka, kisha chapisha orodha.

# --hints--

Unapaswa kuongeza `1` kwenye `my_list` kwa kielezo `1` ukitumia `insert()`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('my_list.insert(1, 1)')
    `)
})
```

`my_list` inapaswa kuwa na vipengele 4.

```js
({
    test: () => runPython(`
    assert len(my_list) == 4
    `)
})
```

`my_list` inapaswa kuwa sawa na `[0, 1, 2, 3]`.

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2, 3]
    `)
})
```

Unapaswa kuchapisha orodha baada ya kuongeza kipengele kipya.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.insert(1, 1)', "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

--fcc-editable-region--
```
