---
id: 615f3b091162165948e1cb33
title: Passaggio 11
challengeType: 0
dashedName: step-11
---

# --description--

Se ispezioni il tuo elemento `.label` con gli strumenti per sviluppatori del browser, potresti notare che in realtà è largo 288 pixel invece di 270. Questo succede perché, per impostazione predefinita, il browser include il bordo e la spaziatura interna quando determina la dimensione di un elemento.

Per risolvere questo problema, resetta il modello del riquadro creando un selettore `*` e impostandogli una proprietà `box-sizing` con valore `border-box`.

# --hints--

Dovresti creare un selettore `*`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*'));
```

Il tuo selettore `*` dovrebbe avere una proprietà `box-sizing` impostata su `border-box`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}
```
