---
id: 68e76b1ce8216249d04df670
title: Criar uma Classe Planet
challengeType: 27
dashedName: build-a-planet-class
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma classe chamada `Planet`.
1. A classe `Planet` deve ter um método `__init__` que:
   - Possui quatro parâmetros: `self`, `name`, `planet_type` e `star`.
   - Lança um `TypeError` com a mensagem `name, planet type, and star must be strings` se algum dos argumentos passados não for do tipo string.
   - Levanta um `ValueError` com a mensagem `name, planet_type, and star must be non-empty strings` se algum dos argumentos passados for uma string vazia.
   - Atribui os valores passados aos atributos da instância `name`, `planet_type` e `star`.
1. A classe `Planet` deve ter um método `orbit` que retorna uma string no formato `{name} is orbiting around {star}...`.
1. A classe `Planet` deve ter um método `__str__` que retorne uma string no formato `Planet: {name} | Type: {planet_type} | Star: {star}`.
1. Você deve criar três instâncias da classe `Planet` chamadas `planet_1`, `planet_2` e `planet_3`.
1. Você deve imprimir cada objeto planet para ver a saída do método `__str__`.
1. Você deve chamar o método `orbit` em cada objeto planeta e imprimir o resultado.

# --hints--

Você deve criar uma classe chamada `Planet`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_class('Planet')
`
    )
})
```

A classe `Planet` deve ter um método `__init__`.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__init__')

`
    )
})
```

O método `__init__` deve ter quatro parâmetros: `self`, `name`, `planet_type` e `star`, nesta ordem.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').find_function('__init__').has_args('self, name, planet_type, star')

`
    )
})
```

O método `__init__` deve lançar um `TypeError` com a mensagem `name, planet type, and star must be strings` se o primeiro argumento passado não for do tipo string.

```js
({
    test: () => runPython(
`
try:
  planet = Planet(8, "planet", "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

O método `__init__` deve lançar um `TypeError` com a mensagem `name, planet type, and star must be strings` se o segundo argumento passado não for do tipo string.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", [], "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

O método `__init__` deve lançar um `TypeError` com a mensagem `name, planet type, and star must be strings` se o terceiro argumento passado não for do tipo string.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", 0)
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

O método `__init__` deve lançar um `ValueError` com a mensagem `name, planet_type, and star must be non-empty strings` se o primeiro argumento passado for uma string vazia.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("", "planet", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

O método `__init__` deve lançar um `ValueError` com a mensagem `name, planet_type, and star must be non-empty strings` se o segundo argumento passado for uma string vazia.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

O método `__init__` deve lançar um `ValueError` com a mensagem `name, planet_type, and star must be non-empty strings` se o terceiro argumento passado for uma string vazia.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", "")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Você deve atribuir `name` a `self.name` no método `__init__`.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.name').is_equivalent('self.name = name')
`
    )
})
```

Você deve atribuir `planet_type` a `self.planet_type` no método `__init__`.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.planet_type').is_equivalent('self.planet_type = planet_type')
`
    )
})
```

Você deve atribuir `star` a `self.star` no método `__init__`.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.star').is_equivalent('self.star = star')
`
    )
})
```

A classe `Planet` deve ter um método `orbit`.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('orbit')
`
    )
})
```

O método `orbit` deve retornar uma string no formato `{name} is orbiting around {star}...`.

```js
({
    test: () => runPython(
`
pl = Planet('planetino', 'type', 'star')
assert pl.orbit() == 'planetino is orbiting around star...'
`
    )
})
```

A classe `Planet` deve ter um método `__str__`.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__str__')
`
    )
})
```

O método `__str__` deve retornar uma string no formato `Planet: {name} | Type: {planet_type} | Star: {star}`.

```js
({
    test: () => runPython(
`
assert str(Planet("abcd", "bcdea", "jhkl")) == "Planet: abcd | Type: bcdea | Star: jhkl"
`
    )
})
```

Você deve criar três instâncias da classe `Planet` chamadas `planet_1`, `planet_2` e `planet_3`.

```js
({
    test: () => runPython(
`
assert isinstance(planet_1, Planet)
assert isinstance(planet_2, Planet)
assert isinstance(planet_3, Planet)
`
    )
})
```

Você deve imprimir cada objeto planet para ver a saída do método `__str__`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1)')
assert _Node(_code).has_call('print(planet_2)')
assert _Node(_code).has_call('print(planet_3)')
`
    )
})
```

Você deve chamar o método `orbit` em cada objeto planeta e imprimir o resultado.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1.orbit())')
assert _Node(_code).has_call('print(planet_2.orbit())')
assert _Node(_code).has_call('print(planet_3.orbit())')
`
    )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class Planet:
    def __init__(self, name, planet_type, star):
        args = (name, planet_type, star)
        if not all(isinstance(arg, str) for arg in args):
            raise TypeError("name, planet type, and star must be strings")
        if not all(args):
            raise ValueError('name, planet_type, and star must be non-empty strings')
        self.name = name
        self.planet_type = planet_type
        self.star = star   # The star the planet orbits

    def orbit(self):
        return f'{self.name} is orbiting around {self.star}...'

    def __str__(self):
        return f'Planet: {self.name} | Type: {self.planet_type} | Star: {self.star}'


# Test the Planet class
planet_1 = Planet('Earth', 'terrestrial', 'Sun')
planet_2 = Planet('Jupiter', 'gas giant', 'Sun')
planet_3 = Planet('Kepler-452b', 'super-Earth', 'Kepler-452')

print(planet_1.orbit()) # Earth is orbiting around Sun...
print(planet_1) # Planet: Earth | Type: terrestrial | Star: Sun

print(planet_2.orbit()) # Jupiter is orbiting around Sun...
print(planet_2) # Planet: Jupiter | Type: gas giant | Star: Sun

print(planet_3.orbit()) # Kepler-452b is orbiting around Kepler-452...
print(planet_3) # Planet: Kepler-452b | Type: super-Earth | Star: Kepler-452
```
