---
id: a77dbc43c33f39daa4429b4f
title: ブール値チェック関数を作成する
challengeType: 26
dashedName: build-a-boolean-check-function
---

# --description--

このラボでは、値がブール値プリミティブとして分類されるかどうかをチェックする関数を作成します。

ブール値プリミティブは `true` と `false` です。

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー:**

1. `booWho` という名前の関数があり、1つの引数を受け取る必要があります。
1. 受け取った引数がブール値プリミティブの場合、関数は `true` を返す必要があります。
1. 引数がその他の値の場合、関数は `false` を返す必要があります。

# --hints--

`booWho` 関数が必要です。

```js
assert.isFunction(booWho);
```

`booWho(true)` は `true` を返すべきです。

```js
assert.isTrue(booWho(true));
```

`booWho(false)` は `true` を返すべきです。

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])` は `false` を返すべきです。

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` は `false` を返すべきです。

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })` は `false` を返すべきです。

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` は `false` を返すべきです。

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` は `false` を返すべきです。

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` は `false` を返すべきです。

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` は `false` を返すべきです。

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` は `false` を返すべきです。

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

```
