---
id: 6ac0ae563f6de70930f47ef7
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

عند كتابة الكود، غالبًا ما تريد إعادة استخدام أو الإشارة إلى قيم موجودة في كودك. المتغيرات تتيح لك ذلك.

لتعيين قيمة إلى متغير، اكتب اسم المتغير متبوعًا بمعامل التعيين (`=`) والقيمة التي تريد تعيينها لذلك المتغير:

```py
sentence = 'I am learning Python'
```

انتقل إلى السطر المميز، أنشئ متغيرًا باسم `sentence` وعيّن له السلسلة النصية `I am learning Python`.

# --hints--

يجب أن يحتوي كودك على `sentence = 'I am learning Python'`. تأكد من البدء في بداية السطر.

```js
({ test: () => runPython(`
import re
target_var = _Node(_code).find_variable("sentence")
pattern = r"""sentence\\s*=\\s*(['"])\\s*i\\s+am\\s+learning\\s+python\\s*\\1\\s*"""
assert re.fullmatch(pattern, str(target_var).strip(), flags=re.IGNORECASE)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')
--fcc-editable-region--

--fcc-editable-region--
```
