---
id: 6ac0a4463662d9738c393e21
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

في الخطوة السابقة، أنشأت **سلسلة نصية**: متتالية من الأحرف محاطة إما بعلامات اقتباس مفردة أو علامات اقتباس مزدوجة. في البرمجة، تُستخدم السلاسل النصية لتمثيل النص.

**الطرفية** هي المكان الذي يمكنك فيه رؤية ناتج الكود الخاص بك.

لكي تتمكن من رؤية الناتج، حوّل السطر الحالي من الكود إلى:

```py
print('Hello')
```

تأكد من البدء من بداية السطر.

# --hints--

يجب أن يحتوي كودك على `print('Hello')`. تأكد من البدء من بداية السطر.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*hello\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
'Hello'
--fcc-editable-region--
```
