---
id: 6ac09f8420b4b67e075d2cd2
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذه الورشة، ستنشئ أول برنامج Python خاص بك.

هذا هو محرر الكود الخاص بك، حيث ستكتب كود Python. السطر 1 مميز. اذهب إليه واكتب:

```py
'Hello'
```

تأكد من البدء في بداية السطر ولا تنسَ علامات الاقتباس.

عندما تنتهي، انقر على زر "التحقق من الكود" لترى ما إذا كان صحيحًا.

# --hints--

يجب أن تكتب `'Hello'`. تأكد من البدء في بداية السطر ولا تنسَ علامات الاقتباس.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""(['"])\\s*hello\\s*\\1"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
