---
id: 6ac0aaba914d81b32567c076
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

`print` دالة مدمجة يمكنك استخدامها لإرسال الإخراج إلى الطرفية. كما ترى، النص `Hello` أصبح ظاهرًا الآن على الطرفية.

الآن، اطبع النص `freeCodeCamp` على الطرفية بالذهاب إلى السطر المميز وكتابة:

```py
print('freeCodeCamp')
```

تأكد من البدء في بداية السطر.

# --hints--

يجب أن يحتوي كودك على `print('freeCodeCamp')`. تأكد من البدء في بداية السطر.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*freecodecamp\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
--fcc-editable-region--

--fcc-editable-region--
```
