---
id: 6ac0b35f3b29a27f7e207295
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

في الخطوة السابقة، أنشأت المتغير `sentence`. الآن اطبع قيمته على الطرفية بكتابة `print(sentence)` في السطر المميز.

بهذا تكون قد أكملت برنامج Python الأول الخاص بك.

# --hints--

يجب أن يحتوي كودك على `print(sentence)`. تأكد من البدء في بداية السطر.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(sentence)")`)) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
print(sentence)
```
