---
id: 6a57c331ae94fa4631b60863
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

博物館には複数のアーティファクトが記録されています。最初のものと同じ形で、キー `102` の下に2つ目のアーティファクトを `collection` に追加してください。

これらの値を使ってください：

```js
{
  title: "Bronze Tablet",
  category: "Inscription",
  curator: {
    id: 202,
    name: "Robert Sinclair",
  },
  locations: [{ gallery: "Archive Wing", year: 2019 }],
  tags: ["bronze", "writing"],
  onDisplay: false,
}
```

# --hints--

`collection` には2つのエントリーがあるはずです。

```js
assert.strictEqual(Object.keys(collection).length, 2);
```

`102` の下に保存されているアーティファクトは、説明されたオブジェクトと一致している必要があります。

```js
assert.deepEqual(collection[102], {
  title: "Bronze Tablet",
  category: "Inscription",
  curator: {
    id: 202,
    name: "Robert Sinclair",
  },
  locations: [{ gallery: "Archive Wing", year: 2019 }],
  tags: ["bronze", "writing"],
  onDisplay: false,
});
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
--fcc-editable-region--
  
--fcc-editable-region--
};
```
