---
id: 6a57c331ae94fa4631b6086c
title: ステップ 13
challengeType: 1
dashedName: step-13
---

# --description--

時間が経つと、アーティファクトは頻繁に表示されたり非表示になったりします。

`id` パラメータを受け取る `toggleDisplayStatus` 関数を作成してください。`collection[id]` を使ってアーティファクトを検索します。存在する場合、その `onDisplay` プロパティを現在の値の反対に設定してください。

# --hints--

`toggleDisplayStatus` という名前の関数があり、`id` というパラメータを持っている必要があります。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.toggleDisplayStatus);
const params = explorer.allFunctions.toggleDisplayStatus.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

`toggleDisplayStatus` を呼び出すと、アーティファクトの `onDisplay` の値が両方向に切り替わるはずです。

```js
const before101 = collection[101].onDisplay;
toggleDisplayStatus(101);
assert.strictEqual(collection[101].onDisplay, !before101);
toggleDisplayStatus(101);
assert.strictEqual(collection[101].onDisplay, before101);
const before102 = collection[102].onDisplay;
toggleDisplayStatus(102);
assert.strictEqual(collection[102].onDisplay, !before102);
toggleDisplayStatus(102);
assert.strictEqual(collection[102].onDisplay, before102);
```

`toggleDisplayStatus` は、`id` に一致するアーティファクトがない場合にエラーをスローしたり、コレクションを変更したりしてはいけません。

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => toggleDisplayStatus(999));
assert.strictEqual(JSON.stringify(collection), before);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

--fcc-editable-region--

--fcc-editable-region--
```
