---
id: 6a57c331ae94fa4631b60861
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

すべてのアーティファクトには責任者がいます。

キーが`101`の下に保存されているアーティファクトに`curator`プロパティを追加し、その値を`id`が`201`で`name`が`"Earl Sinclair"`のオブジェクトに設定してください。

# --hints--

キーが`101`のアーティファクトには`curator`プロパティが必要です。

```js
assert.property(collection[101], "curator");
```

`curator`はオブジェクトである必要があります。

```js
assert.isObject(collection[101].curator);
```

`curator`は`id`が`201`である必要があります。

```js
assert.strictEqual(collection[101].curator.id, 201);
```

`curator` は `name` プロパティを持ち、その値は `"Earl Sinclair"` である必要があります。

```js
assert.strictEqual(collection[101].curator.name, "Earl Sinclair");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
