---
id: 6a57c331ae94fa4631b60872
title: ステップ 19
challengeType: 1
dashedName: step-19
---

# --description--

次に、アーティファクトを要約するテンプレートリテラルを返してください。タイトルを最初の行に置き、その後にカテゴリ、キュレーターの名前、現在のギャラリー、展示状況をそれぞれ1行ずつ続けてください。ギャラリーには `currentLocation.gallery` を使います。

例えば、`buildSummary(101)` を呼び出すと、次のように返す必要があります：

```md
Golden Mask
Category: Ceremonial
Curator: Fran Sinclair
Current Gallery: Hall C
On Display: true
```

# --hints--

`buildSummary(101)` にはアーティファクトのタイトルである `Golden Mask` を含める必要があります。

```js
assert.include(buildSummary(101), "Golden Mask");
```

`buildSummary(101)` には `Category: Ceremonial` の行を含める必要があります。

```js
assert.include(buildSummary(101), "Category: Ceremonial");
```

`buildSummary(101)` には `Curator: Fran Sinclair` の行を含める必要があります。

```js
assert.include(buildSummary(101), "Curator: Fran Sinclair");
```

`buildSummary(101)` には `Current Gallery: Hall C` の行を含め、`currentLocation.gallery` を使う必要があります。

```js
assert.include(buildSummary(101), "Current Gallery: Hall C");
```

`buildSummary(101)` には `On Display: true` の行を含める必要があります。

```js
assert.include(buildSummary(101), "On Display: true");
```

`buildSummary(101)` は各フィールドをそれぞれ別の行に配置する必要があります。

```js
assert.isAtLeast(buildSummary(101).split("\n").length, 5);
```

`buildSummary(102)` は最新の場所を使い、`Current Gallery: Hall B` を表示する必要があります。

```js
assert.include(buildSummary(102), "Current Gallery: Hall B");
```

`buildSummary` はどのアーティファクトの値からでも要約を作成し、各フィールドを1行ずつ順番に並べる必要があります。

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [
    { gallery: "Hall Y", year: 2001 },
    { gallery: "Hall Z", year: 2002 },
  ],
  tags: [],
  onDisplay: false,
};
const summary = buildSummary(457);
assert.deepEqual(
  summary.split("\n").map((line) => line.trim()),
  [
    "Test Artifact",
    "Category: Test Category",
    "Curator: Test Curator",
    "Current Gallery: Hall Z",
    "On Display: false",
  ]
);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

function updateCurator(id, name) {
  const artifact = collection[id];
  if (artifact) {
    artifact.curator.name = name;
  }
}

updateCurator(101, "Fran Sinclair");
console.log(collection[101].curator.name);

function buildSummary(id) {
  const artifact = collection[id];
  if (!artifact) {
    return "Artifact not found";
  }
  const currentLocation = artifact.locations[artifact.locations.length - 1];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
