---
id: 6a57c331ae94fa4631b60868
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

キュレーターはアーティファクトにタグを付けて検索しやすくしますが、タグは二重に追加してはいけません。新しいタグを追加する前に、`includes()` メソッドを使ってすでに存在するかどうかを確認できます。配列に値がすでに含まれている場合は `true` を返します。

`id` と `tag` のパラメータを受け取る `addTag` という名前の関数を作成してください。`collection[id]` を使ってアーティファクトを検索します。存在し、かつその `tags` 配列に `tag` がまだ含まれていなければ、`tag` を `artifact.tags` に追加してください。

# --hints--

`id` と `tag` というパラメータを持つ `addTag` という名前の関数を用意してください。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addTag);
const params = explorer.allFunctions.addTag.parameters;
assert.equal(params.length, 2);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "tag");
```

`addTag` を新しいタグで呼び出すと、そのタグがアーティファクトの `tags` 配列に追加されるはずです。

```js
addTag(102, "ancient");
assert.include(collection[102].tags, "ancient");
```

`addTag` をアーティファクトがすでに持っているタグで呼び出しても、重複して追加されてはいけません。

```js
const tagsBefore = [...collection[102].tags];
addTag(102, "bronze");
assert.deepEqual(collection[102].tags, tagsBefore);

const start = collection[102].tags.length;
addTag(102, "unique-check");
addTag(102, "unique-check");
assert.strictEqual(collection[102].tags.length, start + 1);
```

`addTag` は `id` が指すどのアーティファクトにもタグを追加するはずです。

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
addTag(457, "test-tag");
assert.deepEqual(collection[457].tags, ["test-tag"]);
```

`id` に一致するアーティファクトがない場合、`addTag` はエラーをスローしたり、コレクションを変更したりしてはいけません。

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => addTag(999, "test"));
assert.strictEqual(JSON.stringify(collection), before);
```

`tag` はアーティファクトの既存の `tags` 配列に追加してください。配列を置き換えてはいけません。

```js
collection[458] = {
  title: "Push Test",
  category: "Test Category",
  curator: { id: 458, name: "Test Curator" },
  locations: [{ gallery: "Hall Z", year: 2002 }],
  tags: [],
  onDisplay: false,
};
const tagsBefore = collection[458].tags;
const spy = __helpers.spyOn(tagsBefore, "push");
try {
  addTag(458, "pushed");
  assert.deepEqual(spy.calls, [["pushed"]]);
} finally {
  spy.restore();
}
assert.strictEqual(collection[458].tags, tagsBefore);
assert.deepEqual(collection[458].tags, ["pushed"]);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

--fcc-editable-region--

--fcc-editable-region--
```
