---
id: 6a57c331ae94fa4631b60866
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

時々、スタッフが博物館にないアーティファクトを求めることがあります。もしその`id`に対応するアーティファクトが保存されていなければ、`collection[id]`は`undefined`となり、そこから`.title`を読み取ろうとするとエラーがスローされます。

return文を条件付きに更新してください。アーティファクトが見つかった場合は`artifact.title`を返し、見つからなかった場合は文字列`"Artifact not found"`を返してください。

# --hints--

`getArtifactTitle(101)` は引き続き `"Golden Mask"` を返す必要があります。

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

`getArtifactTitle` は、`id` に一致するアーティファクトがない場合、`"Artifact not found"` を返す必要があります。

```js
assert.strictEqual(getArtifactTitle(999), "Artifact not found");
```

`getArtifactTitle`は、`101`だけでなく見つけたどのアーティファクトのタイトルも返すべきです。

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
assert.strictEqual(getArtifactTitle(457), "Test Artifact");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
--fcc-editable-region--
  return artifact.title;
--fcc-editable-region--
}
```
