---
id: 6a57c331ae94fa4631b60864
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

監査は、すでに記録されている内容を読むことから始まります。各アーティファクトはIDの下にストアされているので、そのIDを大括弧で囲むことでアクセスできます。例えば `collection[101]` のようにします。

`console.log()` を `collection[101].title` と一緒に呼び出して、最初のアーティファクトのタイトルをログに出してください。次に、`console.log()` を `collection[101].curator.name` と一緒に呼び出して、そのキュレーターの名前をログに出してください。

# --hints--

最初のアーティファクトの `title` をコンソールにログ出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*title|\[\s*["']title["']\s*\])\s*\)/);
```

最初のアーティファクトのキュレーターの `name` をコンソールにログ出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

まず `title` をログ出力し、次にキュレーターの `name` をログ出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*title|\[\s*["']title["']\s*\])\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

--fcc-editable-region--

--fcc-editable-region--
```
