---
id: 6a57c331ae94fa4631b60862
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

各アーティファクトは、どこにあったか、どのように説明されているか、そして展示されているかどうかも追跡しています。

`101` に保存されているアーティファクトに、さらに3つのプロパティを追加してください。

- `locations`: 2つのオブジェクト `{ gallery: "Hall A", year: 2020 }` と `{ gallery: "Hall C", year: 2024 }` を含む配列
- `tags`: 文字列 `"gold"` と `"egypt"` を含む配列
- `onDisplay`: ブール値の `true`

# --hints--

アーティファクトの `locations` は、指定された2つの場所オブジェクトの配列である必要があります。

```js
assert.deepEqual(collection[101].locations, [
  { gallery: "Hall A", year: 2020 },
  { gallery: "Hall C", year: 2024 },
]);
```

アーティファクトの `tags` は `["gold", "egypt"]` である必要があります。

```js
assert.deepEqual(collection[101].tags, ["gold", "egypt"]);
```

アーティファクトの `onDisplay` は `true` である必要があります。

```js
assert.strictEqual(collection[101].onDisplay, true);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
