---
id: 6a57c331ae94fa4631b60865
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

スタッフは1日に何度もアーティファクトを調べるので、監査担当者には再利用可能な検索機能が必要です。大括弧は変数も受け取れるので、`collection[id]` はその `id` が指すアーティファクトを返します。

`id` パラメータを受け取る `getArtifactTitle` という名前の関数を作成してください。`collection[id]` を使ってアーティファクトを取得し、`artifact` という変数にストアしてから、`artifact.title` を返してください。

# --hints--

`id` というパラメータを持つ `getArtifactTitle` という名前の関数が必要です。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle);
const params = explorer.allFunctions.getArtifactTitle.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

`collection[id]` を `artifact` という変数に代入してください。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle?.variables.artifact);
assert.isTrue(explorer.allFunctions.getArtifactTitle?.variables.artifact?.value.matches("collection[id]"));
```

`getArtifactTitle` は `artifact.title` を返す必要があります。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getArtifactTitle?.hasReturn("artifact.title"));
```

`getArtifactTitle(101)` は `"Golden Mask"` を返すべきです。

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

`getArtifactTitle(102)` は `"Bronze Tablet"` を返すべきです。

```js
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

--fcc-editable-region--

--fcc-editable-region--
```
