---
id: 6a57c331ae94fa4631b6086b
title: ステップ 12
challengeType: 1
dashedName: step-12
---

# --description--

`id` が `102` のアーティファクトを `2026` 年に `"Hall B"` に移動してください。その後、`console.log()` を呼び出して `collection[102].locations` をログに出力し、そのアーティファクトの locations 配列を表示してください。

# --hints--

`moveArtifact`を`102`、`"Hall B"`、および`2026`で呼び出してください。

```js
assert.match(__helpers.removeJSComments(code), /moveArtifact\s*\(\s*102\s*,\s*["']Hall B["']\s*,\s*2026\s*\)/);
```

2番目のアーティファクトの`locations`をコンソールにログしてください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?102["']?\s*\](?:\.\s*locations|\[\s*["']locations["']\s*\])\s*\)/);
```

アーティファクトを移動した後に、`locations` をログに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /moveArtifact\s*\(\s*102\s*,\s*["']Hall B["']\s*,\s*2026\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?102["']?\s*\](?:\.\s*locations|\[\s*["']locations["']\s*\])\s*\)/);
```

2番目のアーティファクトの最新の場所は `{ gallery: "Hall B", year: 2026 }` である必要があります。

```js
assert.deepEqual(collection[102].locations[collection[102].locations.length - 1], { gallery: "Hall B", year: 2026 });
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
