---
id: 6a57c331ae94fa4631b6086e
title: ステップ 15
challengeType: 1
dashedName: step-15
---

# --description--

博物館が成長するにつれて、遺物の担当者が変わることがあります。

`id` と `name` をパラメータに取る `updateCurator` 関数を作成してください。`collection[id]` を使って遺物を調べます。存在する場合は、その遺物の `curator` オブジェクトの `name` プロパティを指定された `name` に更新してください。

# --hints--

`updateCurator` という名前の関数を用意し、パラメータ名は `id` と `name` にしてください。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.updateCurator);
const params = explorer.allFunctions.updateCurator.parameters;
assert.equal(params.length, 2);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "name");
```

`updateCurator` を呼び出すと、遺物の `curator` オブジェクトの `name` が変更されるはずです。

```js
const curatorBefore = { ...collection[102].curator };
updateCurator(102, "Test Curator");
assert.strictEqual(collection[102].curator.name, "Test Curator");
assert.deepEqual(collection[102].curator, {
  ...curatorBefore,
  name: "Test Curator",
});
```

`updateCurator` は、`id` が指す遺物の担当者を変更する必要があります。

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
updateCurator(457, "Fresh Curator");
assert.strictEqual(collection[457].curator.name, "Fresh Curator");
```

`updateCurator` は、`id` に一致する遺物がない場合にエラーをスローしたり、コレクションを変更したりしてはいけません。

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => updateCurator(999, "Test Curator"));
assert.strictEqual(JSON.stringify(collection), before);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

--fcc-editable-region--

--fcc-editable-region--
```
