---
id: 6a57c265abbc95dfebc05098
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、アーティファクトの出所監査ツールを作ります。博物館のアーティファクトをオブジェクトにストアし、それらを検索したり、タグ付けしたり、ギャラリー間で移動したり、履歴をまとめる関数を書きます。

各アーティファクトはIDの下にストアされます。`collection`という名前のオブジェクトを作成し、エントリーを追加してください。キーを`101`に設定し、値を`title`が`"Golden Mask"`で`category`が`"Ceremonial"`のオブジェクトにしてください。

# --hints--

`collection`という名前の変数を宣言してください。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.collection);
```

`collection`はオブジェクトである必要があります。

```js
assert.isObject(collection);
```

`collection`には1つのエントリーが必要です。

```js
assert.strictEqual(Object.keys(collection).length, 1);
```

`101`の下にストアされたアーティファクトはオブジェクトである必要があります。

```js
assert.isObject(collection[101]);
```

`101`の下にストアされたアーティファクトは`title`が`"Golden Mask"`である必要があります。

```js
assert.strictEqual(collection[101].title, "Golden Mask");
```

`101`の下にストアされたアーティファクトは`category`が`"Ceremonial"`である必要があります。

```js
assert.strictEqual(collection[101].category, "Ceremonial");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
