---
id: 61fd6ab779390f49148773bb
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

أسفل عنصر `h1` الخاص بك، أنشئ عنصر `div`. أعطه خاصية `id` مضبوطة على `years`. تريد إخفاء هذا العنصر تحديدًا عن قارئ الشاشة، لذا أعطه خاصية `aria-hidden` مضبوطة على `true`.

# --hints--

يجب أن تنشئ عنصر `div` جديدًا بعد عنصر `h1` الخاص بك.

```js
assert(document.querySelector('h1')?.nextElementSibling?.localName === 'div');
```

يجب أن يحتوي عنصر `div` الجديد على خاصية `id` مضبوطة على `years`.

```js
assert(document.querySelector('div')?.getAttribute('id') === 'years');
```

يجب أن يحتوي عنصر `div` الجديد على خاصية `aria-hidden` مضبوطة على `true`.

```js
assert(document.querySelector('div')?.getAttribute('aria-hidden') === 'true');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
