---
id: 61fd719788899952e67692b9
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

تُستخدم عناصر `thead` و `tbody` لتحديد أي جزء من جدولك هو القسم العلوي، وأي جزء يحتوي على البيانات أو المحتوى الأساسي.

أضف عنصر `thead` و `tbody` إلى أول `table` لديك، أسفل عنصر `caption`.

# --hints--

يجب أن يحتوي أول عنصر `table` لديك على عنصر `thead`.

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('thead'));
```

يجب أن يحتوي أول عنصر `table` لديك على عنصر `tbody`.

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('tbody'));
```

يجب أن يكون عنصر `thead` الخاص بك مباشرة أسفل عنصر `caption`.

```js
assert(document.querySelector('caption')?.nextElementSibling?.localName === 'thead');
```

يجب أن يكون عنصر `tbody` الخاص بك مباشرة أسفل عنصر `thead`.

```js
assert(document.querySelector('thead')?.nextElementSibling?.localName === 'tbody');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
            <caption>Assets</caption>
          </table>
--fcc-editable-region--
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
