---
id: 61fd5a93fd62bb35968adeab
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

أعد إعداد HTML الخاص بك مع `DOCTYPE`، و`html` التي تشير إلى أن هذه الوثيقة باللغة الإنجليزية، ووسوم `head` و`body`.

أضف إلى وسم `head` الخاص بك الوسوم المناسبة من نوع `meta` لـ `charset` و`viewport`، وعنصر `title` بعنوان مناسب، وعنصر `link` لورقة الأنماط الخاصة بك.

# --hints--

يجب أن يحتوي شفرتك على إعلان `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

يجب أن تحتوي شفرتك على وسم `html`.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

يجب أن يحتوي وسم `html` على وسم بداية مع خاصية `lang` بقيمة `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي وسم `html` على وسم الإغلاق.

```js
assert(code.match(/<\/html\s*>/));
```

يجب أن يحتوي شفرتك على وسم `head` داخل وسم `html`.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

يجب أن يحتوي شفرتك على وسم `body` داخل وسم `html`.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

يجب أن يأتي وسم `head` قبل وسم `body`.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

يجب أن يكون لديك وسمان من نوع `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

يجب أن يحتوي أحد وسوم `meta` على خاصية `name` بقيمة `viewport`، و`content` بقيمة `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

يجب أن يحتوي وسم `meta` الآخر على خاصية `charset` بقيمة `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

يجب أن يحتوي شفرتك على وسم `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

يجب أن يحتوي وسم `title` على بعض النصوص.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

يجب أن يحتوي شفرتك على وسم `link`.

```js
assert(/<link/.test(code))
```

يجب أن يكون وسم `link` داخل وسم `head`.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i))
```

يجب أن يحتوي وسم `link` على خاصية `rel` بقيمة `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

يجب أن يحتوي وسم `link` على خاصية `href` بقيمة `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
