---
id: 61fd78621573aa5e8b512f5e
title: الخطوة 15
challengeType: 0
dashedName: step-15
---

# --description--

في عنصر `tr` الثالث، أضف عنصر `th` يحتوي على النص `Savings Funds set aside for emergencies.`. لف هذا النص، ما عدا `Savings `، داخل عنصر `span` مع تعيين الخاصية `class` إلى `description`.

بعد ذلك، أضف ثلاثة عناصر `td` بالنصوص التالية (بالترتيب): `$500`، `$650`، `$728`. اجعل الخاصية `class` في عنصر `td` الثالث تساوي `current`.

# --hints--

يجب أن يحتوي عنصر `tr` الثالث على عنصر `th`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[2]?.querySelector('th'));
```

يجب أن يحتوي عنصر `th` على النص `Savings Funds set aside for emergencies.`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[2]?.querySelector('th')?.innerText === 'Savings Funds set aside for emergencies.');
```

يجب أن تلف النص `Funds set aside for emergencies.` داخل عنصر `span`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[2]?.querySelector('th > span')?.textContent === 'Funds set aside for emergencies.');
```

يجب أن يحتوي عنصر `span` على الخاصية `class` المعينة إلى `description`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[2]?.querySelector('th > span')?.classList?.contains('description'));
```

يجب أن يكون لديك ثلاثة عناصر `td`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[2]?.querySelectorAll('td').length === 3);
```

يجب أن يحتوي عنصر `td` الأول على النص `$500`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[2]?.querySelectorAll('td')?.[0]?.textContent === '$500');
```

يجب أن يحتوي عنصر `td` الثاني على النص `$650`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[2]?.querySelectorAll('td')?.[1]?.textContent === '$650');
```

يجب أن يحتوي عنصر `td` الثالث على النص `$728`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[2]?.querySelectorAll('td')?.[2]?.textContent === '$728');
```

يجب أن تحتوي الخاصية `class` في عنصر `td` الثالث على القيمة `current`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[2]?.querySelectorAll('td')?.[2]?.classList?.contains('current'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="total">
              </tr>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
