---
id: 613297a923965e0703b64796
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Huenda tayari unafahamu kipengele cha `meta`; hutumika kufafanua taarifa kuhusu ukurasa, kama vile kichwa, maelezo, maneno muhimu, na mwandishi.

Toa ukurasa wako kipengele cha `meta` chenye thamani inayofaa ya `charset`.

Sifa ya kipengele ya `charset` inaeleza usimbaji wa herufi wa ukurasa, na, siku hizi, `UTF-8` ndiyo usimbaji pekee unaoungwa mkono na vivinjari vingi.

# --hints--

Unapaswa kuunda kipengele cha `meta` ndani ya kipengele cha `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

Unapaswa kumpa lebo ya `meta` sifa ya kipengele ya `charset` yenye thamani `UTF-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    --fcc-editable-region--
    
    --fcc-editable-region--
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
