---
id: 61329b210dac0b08047fd6ab
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Kuendelea na vipengele vya `meta`, ufafanuzi wa `viewport` unaeleza kivinjari jinsi ya kuonyesha ukurasa. Kujumuisha moja huongeza ufikikaji wa kuona kwenye simu za mkononi, na kuboresha _SEO_ (ubora wa injini za utafutaji).

Ongeza ufafanuzi wa `viewport` wenye sifa ya kipengele ya `content` inayofafanua `width` na `initial-scale` za ukurasa.

# --hints--

Unapaswa kuunda kipengele kingine cha `meta` ndani ya `head`.

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

Unapaswa kumpa `meta` sifa ya kipengele ya `name` yenye thamani ya `viewport`.

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

Unapaswa kumpa `meta` sifa ya kipengele ya `content` yenye thamani ya `width=device-width, initial-scale=1`.

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    --fcc-editable-region--
    
    --fcc-editable-region--
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
