---
id: 6133d11ef548f51f876149e3
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Uvinjari ni sehemu muhimu ya Ufikikaji, na visomaji vya skrini vinategemea wewe kutoa muundo wa ukurasa wako. Hii hufanikishwa kwa kutumia vipengele vya semantic vya HTML.

Ongeza kipengele cha `header` na `main` kwenye ukurasa wako.

Kipengele cha `header` kitatumika kuanzisha ukurasa, pamoja na kutoa menyu ya uvinjari.

Kipengele cha `main` kitakuwa na maudhui kuu ya ukurasa wako.

# --hints--

Unapaswa kuongeza kipengele cha `header` ndani ya `body`.

```js
assert.exists(document.querySelector('body > header'));
```

Unapaswa kuongeza kipengele cha `main` ndani ya `body`.

```js
assert.exists(document.querySelector('body > main'));
```

Kipengele cha `header` kinapaswa kuja kabla ya kipengele cha `main`.

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
