---
id: 613e275749ebd008e74bb62e
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Kwa sasa, `img` inachukua ukubwa wake wa kawaida, ambao ni mkubwa sana. CSS ina kitendakazi cha `max` kinachorejesha kubwa zaidi kati ya seti ya thamani zilizoandikwa kwa koma. Kwa mfano:

```css
img {
  width: max(250px, 25vw);
}
```

Katika mfano hapo juu, upana wa picha utakuwa 250px ikiwa upana wa eneo la kuona ni chini ya pikseli 1000. Ikiwa upana wa eneo la kuona ni zaidi ya pikseli 1000, upana wa picha utakuwa 25vw. Hii ni kwa sababu 25vw ni sawa na asilimia 25 ya upana wa eneo la kuona.

Punguza picha kwa kutumia `id` kama kichaguzi, na kuweka `width` kuwa kubwa zaidi kati ya `10rem` au `18vw`.

# --hints--

Unapaswa kutumia kichaguzi cha `#logo` kulenga kipengele cha `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

Unapaswa kumpa `img` `width` ya `max(10rem, 18vw)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
