---
id: 6143931a113bb80c45546287
title: Hatua 23
challengeType: 0
dashedName: step-23
---

# --description--

Ndani ya kipengele cha pili cha `section`, ongeza vipengele viwili vya `div` vyenye darasa la vitu la `question-block`.

Kisha, ndani ya kila kipengele cha `div.question-block`, ongeza kipengele kimoja cha `h3` chenye maandishi ya nambari zinazoongezeka, kuanzia `1`, na kipengele kimoja cha `fieldset` chenye darasa la vitu la `question`.

# --hints--

Unapaswa kupanga ndani vipengele viwili vya `div` ndani ya kipengele cha pili cha `section`.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div')?.length, 2);
```

Unapaswa kumpa kipengele cha kwanza kipya cha `div` darasa la vitu la `question-block`.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div')?.[0]?.className, 'question-block');
```

Unapaswa kumpa kipengele cha pili kipya cha `div` darasa la vitu la `question-block`.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div')?.[1]?.className, 'question-block');
```

Unapaswa kupanga ndani kipengele kimoja cha `h3` ndani ya kila kipengele cha `div.question-block`.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > h3')?.length, 2);
```

Unapaswa kumpa kipengele cha kwanza cha `h3` maandishi ya `1`.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > h3')[0]?.textContent.trim(), '1');
```

Unapaswa kumpa kipengele cha pili cha `h3` maandishi ya `2`.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > h3')[1]?.textContent.trim(), '2');
```

Unapaswa kupanga ndani kipengele kimoja cha `fieldset` ndani ya kila kipengele cha `div.question-block`.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > fieldset')?.length, 2);
```

Unapaswa kuweka kipengele cha kwanza cha `fieldset` baada ya kipengele cha kwanza cha `h3`.

```js
assert.exists(document.querySelector('section:nth-of-type(2) > div.question-block > h3 + fieldset'));
```

Unapaswa kuweka kipengele cha pili cha `fieldset` baada ya kipengele cha pili cha `h3`.

```js
assert.exists(document.querySelector('section:nth-of-type(2) > div.question-block:nth-of-type(2) > h3 + fieldset'));
```

Unapaswa kumpa kipengele cha kwanza cha `fieldset` darasa la vitu la `question`.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > fieldset')?.[0]?.className, 'question');
```

Unapaswa kumpa kipengele cha pili cha `fieldset` darasa la vitu la `question`.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > fieldset')?.[1]?.className, 'question');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a href="#student-info">INFO</a></li>
          <li><a href="#html-questions">HTML</a></li>
          <li><a href="#css-questions">CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">Student Info</h2>
          <div class="info">
            <label for="student-name">Name:</label>
            <input type="text" name="student-name" id="student-name" />
          </div>
          <div class="info">
            <label for="student-email">Email:</label>
            <input type="email" name="student-email" id="student-email" />
          </div>
          <div class="info">
            <label for="birth-date">Date of Birth:</label>
            <input type="date" name="birth-date" id="birth-date" />
          </div>
        </section>
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
          --fcc-editable-region--
          
          --fcc-editable-region--
        </section>
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

h1,
h2 {
  font-family: Verdana, Tahoma;
}

h2 {
  border-bottom: 4px solid #dfdfe2;
}


```
