---
id: 614ccc21ea91ef1736b9b578
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Kwa warsha hii, umepewa muundo wa msingi wa msingi.

Anza safari hii ya ufikikaji kwa kutoa sifa ya kipengele ya `lang` kwa kipengele chako cha `html`. Hii itasaidia kisomaji cha skrini kutambua lugha ya ukurasa.

Pia, kumbuka kuunganisha karatasi yako ya mitindo na ukurasa.

# --hints--

Unapaswa kumpa kipengele cha `html` sifa ya kipengele ya `lang`. _Kidokezo: Unaweza kutumia thamani ya `en` kwa Kiingereza._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

Unapaswa kuwa na kipengele cha `link` ndani ya kipengele chako cha `head`.

```js
assert.exists(document.querySelector("head>link"));
```

Kipengele chako cha `link` kinapaswa kuwa na `rel="stylesheet"` na `href="styles.css"`.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>

  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
