---
id: 690fe8b8e39ebd4a4d217d6b
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 다양한 버그 이모지를 표시하는 버그 종 선택기를 만들면서 TypeScript 추상 클래스와 제네릭에 대해 배우세요.

프로젝트에 필요한 HTML과 CSS는 제공되었으니 TypeScript에 집중할 수 있습니다. 나중에 선택할 요소에 대해 알아보려면 파일을 확인하세요.

먼저 모든 버그의 추상 기본 클래스를 만들어야 합니다. 추상 클래스는 직접 인스턴스화할 수 없으며 다른 클래스가 확장하도록 설계되었습니다.

`Bug`라는 이름의 제네릭 타입 매개변수 `T`를 가진 추상 클래스를 만드세요.

# --hints--

`Bug`라는 추상 클래스가 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /abstract\s+class\s+Bug/);
```

`Bug` 클래스는 제네릭 타입 매개변수 `T`를 가져야 합니다.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.classes.Bug.typeParameters[0].matches("T"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bug Emoji Picker</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <form id="bug-form">
        <label for="species">Species:</Label>
        <select id="species">
          <option value="bee" selected>🐝 Bee</option>
          <option value="spider">🕷️ Spider</option>
        </select>
      </form>
      <div>
        <p id="bug-emoji">🐝</p>
      </div>
    </main>
  </body>
  <script src="index.ts"></script>
</html>
```

```css
:root {
  --gradient-back: #987284;
  --gradient-left: #F9B5AC;
  --gradient-right: #AA7674;
  --card: #FFFFFF;
  --card-shadow: 0 10px 30px rgba(30, 41, 59, 0.12);
  --font: #222;
}

html {
  width: 100%;
  min-height: 100%;
  background: radial-gradient(100% 100% at 10% 10%, var(--gradient-left), transparent 60%),
              radial-gradient(90% 85% at 90% 0%, var(--gradient-right), transparent 55%),
              linear-gradient(180deg, var(--gradient-back), #ffffff);
}

body {
  margin: 5rem auto;
  padding: 2rem 3rem;
  width: fit-content;
  min-width: 15%;
  color: var(--font);
  font: 500 16px/1.4 system-ui;
  background: rgba(255, 255, 255, 0.3);
  border: 1px solid rgba(255, 255, 255, 0.3);
  border-radius: 15%;
  box-shadow: var(--card-shadow);

  backdrop-filter: blue(12px);
  -webkit-backdrop-filter: blur(12px);
}

select, input {
  appearance: none;
  -webkit-appearance: none;
  -moz-appearance: none;

  background: rgba(255, 255, 255, 0.25);
  border-radius: 12px;
  padding: 0.1rem 2rem 0.1rem 0.8rem;
  font-size: 1rem;
  box-shadow: 0 4px 12px rgba(0, 0, 0, 0.15);
}

form { 
  display: flex;
  flex-direction: column;
}

#bug-emoji {
  margin: 0 auto;
  width: fit-content;
  font-size: 128px;
}

#bug-name {
  margin: 0 auto;
  width: fit-content;
  font-size: 24;
  font-weight: 700;
}
```

```ts
--fcc-editable-region--

--fcc-editable-region--
```
