---
id: 66bb3e20d3dc5b6d0a21f5dd
title: Jenga kielekezi cha wakati halisi
challengeType: 25
dashedName: build-a-real-time-counter
demoType: onClick
---

# --description--

Katika maabara hii, utajenga kielekezi cha herufi cha wakati halisi kwa kutumia JavaScript. Kielekezi cha herufi kitaonyesha idadi ya herufi zilizoingizwa katika kipengele cha `textarea`. Kielekezi kitaendelea kusasishwa kwa wakati halisi wakati mtumizi anapochapa katika `textarea`.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara hii.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kipengele cha `textarea` chenye sifa ya kipengele cha `id` ya `text-input`.
2. Inapaswa kuwepo na kipengele cha `p` chenye sifa ya kipengele cha `id` ya `char-count`. Kwanza kitakuwa na maandishi ya `Character Count: 0/50`. Maandishi haya ya kiwakilishi yanapaswa kubadilishwa na kuonyesha hesabu ya sasa ya herufi katika kipengele cha `textarea`.
3. Wakati kipengele cha `#text-input` kina maandishi ya `hello coder` kipengele cha `#char-count` kinapaswa kuwa na maandishi ya `"Character Count: 11/50"`.
4. Wakati kipengele cha `#text-input` kina maandishi ya `i am a programmer` kipengele cha `#char-count` kinapaswa kuwa na maandishi ya `"Character Count: 17/50"`.
5. Wakati kipengele cha `#text-input` kina maandishi ya `hello world` kipengele cha `#char-count` kinapaswa kuwa na maandishi ya `"Character Count: 11/50"`.
6. Wakati kipengele cha `#text-input` kina maandishi ya `I am learning a new language and it is called c++.` kipengele cha `#char-count` kinapaswa kuwa na maandishi ya `"Character Count: 50/50"`. **KUMBUKA**: Ingawa sifa ya `maxlength` ingeweza kufanikisha kazi hii ipasavyo, maabara hii inahitaji suluhisho la JavaScript pekee.  
7. Mtumizi haapaswi kuweza kuingiza zaidi ya herufi `50` katika kipengele cha `textarea`. Wakati hesabu ya herufi itakapofikia `50`, ingizo lolote ziada linapaswa kukatwa moja kwa moja na maandishi ya `Character Count: 50/50` yanaonyeshwa katika `red`.

**Kumbuka:** Hakikisha umeunganisha faili lako la mtindo na faili la JavaScript katika HTML yako.

# --hints--

Unapaswa kuwa na kipengele cha `textarea` chenye sifa ya kipengele cha `id` ya `text-input`. 

```js
const textInput = document.querySelector('textarea');
assert.equal(textInput.id, 'text-input');
```

Unapaswa kuwa na kipengele cha `p` chenye sifa ya kipengele cha `id` ya `char-count`.

```js
assert.exists("p#char-count");
```

Kipengele cha `#char-count` kinapaswa kuwa na maandishi ya `Character Count: 0/50` mwanzoni.

```js
const charCount = document.querySelector('p#char-count')?.innerText;
assert.equal(charCount, 'Character Count: 0/50');
```

Wakati kipengele cha `#text-input` kina maandishi ya `hello coder` kipengele cha `#char-count` kinapaswa kuwa na maandishi ya `"Character Count: 11/50"`

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "hello coder";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(charCount.innerText.trim(), 'Character Count: 11/50');
```

Wakati hesabu ya herufi ni `50`, maandishi yanapaswa kuonyeshwa katika `red`.

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "I am learning a new language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
const computedStyle = window.getComputedStyle(charCount);
assert.oneOf(computedStyle.color, [
  "red",
  "rgb(255, 0, 0)",
  "rgb(100%, 0%, 0%)",
  "rgba(255, 0, 0, 1)",
  "rgba(100%, 0%, 0%, 1)",
  "hsl(0, 100%, 50%)",
  "hsla(0, 100%, 50%, 1)"
]);
```

Ikiwa hesabu ya herufi ni kubwa au sawa na `50`, ingizo ziada zinapaswa kukatwa.

```js
const textInput = document.getElementById('text-input');
textInput.value = "I am learning a new programming language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(textInput.value.length, 50);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real Time Counter</title>
</head>

<body>

</body>


</html>
```

```css

```

```js
    
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real-time Character Counter</title>
    <link rel="stylesheet" href="styles.css">

<body>
    <div class="container">
        <h1>Real-Time Character Counter</h1>
        <textarea id="text-input" placeholder="Type something..."></textarea>
        <p id="char-count">Character Count: 0/50</p>
    </div>

<script src="script.js"></script>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    font-family: Arial, sans-serif;
    background-color: #c0c1ec;
    margin: 0;
}

h1{
    background-color: #f2f48b;
    padding-left: 10px;
    padding-right: 10px;
}
.container {
    text-align: center;
}

textarea {
    width: 300px;
    height: 100px;
    font-size: 18px;
    padding: 10px;
    border: 1px solid #ccc;
    border-radius: 5px;
    resize: none;
    display: block;
    margin: 0 auto;
    border: 1px solid slategray;
    box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

textarea:focus {
    outline: none;
    border-color: #f2f48b;
}

#char-count {
    margin-top: 10px;
    font-size: 18px;
    color: darkslategray;
    font-weight: bold;
}
```

```js
document.getElementById("text-input").addEventListener("input", function () {
  let textInput = document.getElementById("text-input");
  let charCount = textInput.value.length;

  if (charCount > 50) {
    textInput.value = textInput.value.substring(0, 50);
    charCount = 50;
  }

  let charCountDisplay = document.getElementById("char-count");
  charCountDisplay.innerHTML = `Character Count: ${charCount}/50`;

  if (charCount === 50) {
    charCountDisplay.style.color = "red";
  } 

});
```
