---
id: 68af272f60e009597dd861e1
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個工作坊中，你將創建一個媒體目錄。首先定義一個名為 `Movie` 的類別。

為你的類別新增一個帶有以下參數的 `__init__` 方法：`self`、`title`、`year`、`director` 和 `duration`。在 `__init__` 方法內，將 `title`、`year`、`director` 和 `duration` 參數指定給同名的實體屬性。

# --hints--

你應該創建一個名為 `Movie` 的類別。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Movie")`)
})
```

你的 `Movie` 類別應該有一個 `__init__` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__init__")`)
})
```

你的 `__init__` 方法應該有以下參數：`self`、`title`、`year`、`director` 和 `duration`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_args("self, title, year, director, duration")`)
})
```

你應該在你的 `__init__` 方法中將 `title` 指定給 `self.title`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.title = title")`)
})
```

你應該在你的 `__init__` 方法中將 `year` 指定給 `self.year`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.year = year")`)
})
```

你應該在你的 `__init__` 方法中將 `director` 指定到 `self.director`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.director = director")`)
})
```

你應該在你的 `__init__` 方法中將 `duration` 指定給 `self.duration`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.duration = duration")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
