---
id: 68af273960e009597dd861e2
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

你想要一個可讀的顯示來進行偵錯和列印。創建一個 `__str__` 方法，該方法傳回一個字串，格式為 `title (year) - duration min, director`（其中 `title`、`year`、`duration` 和 `director` 應該被對應屬性的值所替換），以便列印電影時能給出一個漂亮的摘要。

# --hints--

你的 `Movie` 類別應該有一個 `__str__` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__str__")`)
})
```

你的 `__str__` 方法應該傳回格式為 `title (year) - duration min, director` 的字串，其中 `title`、`year`、`duration` 和 `director` 應該取代為對應屬性的值。

```js
({
    test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    print(movie1)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "The Matrix (1999) - 136 min, The Wachowskis" in output
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    
--fcc-editable-region--
```
