---
id: 68af274560e009597dd861e3
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

現在，創建一個 `Movie` 的實體，傳入你選擇的引數並將其指定到名為 `movie1` 的變數。然後，將 `movie1` 列印到主控台。

# --hints--

你應該有一個名為 `movie1` 的變數。

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("movie1")`)
})
```

你應該將一個 `Movie` 的實體指定給 `movie1` 變數。

```js
({
    test: () => runPython(`
    assert isinstance(movie1, Movie)
    `)
})
```

你應該將 `movie1` 列印到主控台。

```js
({
    test: () => runPython(`assert _Node(_code).has_call("print(movie1)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--

--fcc-editable-region--
```
