---
id: 68af277560e009597dd861e8
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

現在刪除你的 `print(movie1)` 行。然後，將你的 `movie1` 指派放入 `try` 區段，並創建一個 `except` 語句來捕捉 `ValueError as e`，並使用 f-string 列印 `Validation Error:`，後面接一個空格和錯誤訊息。

隨意更改傳遞給實例化 `movie1` 的引數，以查看不同的錯誤訊息。

# --hints--

你的程式碼中不應該有 `print(movie1)`。

```js
({
  test: () => runPython(`assert not _Node(_code).has_call("print(movie1)")`)
})
```

你應該有一個包含 `movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)` 的 `try` 區塊。

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_body().has_stmt("movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)")`)
})
```

你應該有一個 `except` 語句來捕捉 `ValueError as e`。

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].has_except("ValueError", "e")`)
})
```

你應該在 `except` 區塊中使用 f-string 來列印 `Validation Error:`，後面接一個空格和錯誤訊息。

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_except("ValueError", "e").has_call("print(f'Validation Error: {e}')")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
--fcc-editable-region--
```
