---
id: 68af275860e009597dd861e5
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

第一部電影是在 1895 年製作的。創建另一個 `if` 述語，檢查 `year` 是否小於 `1895`，並引發帶有訊息 `Year must be 1895 or later` 的 `ValueError`。

# --hints--

當 `year` 小於 `1895` 時，你應該引發一個帶有訊息 `Year must be 1895 or later` 的 `ValueError`。

```js
({
  test: () => runPython(`
  _msg = 'Year must be 1895 or later'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie('The Matrix', 1894, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('The Matrix', 0, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

當 `year` 是有效數字時，你不應該引發 `ValueError`。

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1895, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid year'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
