---
id: 68af274e60e009597dd861e4
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

在前面的步驟中，你創建了一個基本的電影物件。現在，你將添加一些驗證，以確保電影無法使用無效的屬性進行具現化。

在 `__init__` 方法開始時，創建一個 `if` 述語來檢查 `title` 是否為空或僅包含空白字元（你可以使用 `strip` 方法來建構此條件）。

如果是這樣，請引發一個帶有訊息 `Title cannot be empty` 的 `ValueError`。

# --hints--

當 `title` 為空或僅包含空白時，你應該引發一個帶有訊息 `Title cannot be empty` 的 `ValueError`。

```js
({
  test: () => runPython(`
  _msg = 'Title cannot be empty'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie(' ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('   ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

當 `title` 是有效的字串時，你不應該引發 `ValueError`。

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid title'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
