---
id: 68b704862dd6559e890c2a86
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

在 `if` 述語之後，創建一個名為 `result` 的變數。這將是從該方法傳回的字串。

所以指定一個格式為 `Media Catalogue (N items):` 的 f-string。將 `N` 取代為 `items` 列表中的項目數量，並使字串以兩個換行字元結尾。

# --hints--

你應該在你的 `__str__` 方法中有一個名為 `result` 的變數。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("__str__").has_variable("result")`)
})
```

你應該將 `f'Media Catalogue ({len(self.items)} items):\n\n'` 指定給 `result` 變數。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("__str__").has_stmt("result = f'Media Catalogue ({len(self.items)} items):\\\\n\\\\n'")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)
--fcc-editable-region--
    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'
        
--fcc-editable-region--
try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')
```
