---
id: 68b804f3be853a355f055fef
title: 步驟 18
challengeType: 20
dashedName: step-18
---

# --description--

`super` 函式用來取用父類別而不需明確命名它。舉例來說，`super().__init__(arg1, arg2, ...)` 會呼叫父類別中定義的 `__init__` 方法。

目前，你已經有效地覆寫了從 `Movie` 繼承的 `__init__` 方法。這表示 `TVSeries` 物件將只有 `seasons` 屬性和 `total_episodes` 屬性。這不是你想要的。

為了保留所有在 `Movie` 類別的 `__init__` 中定義的程式碼，使用 `super` 來呼叫父類別的 `__init__`。請注意，你需要傳入該方法所需的引數。在這個例子中，為 `title`、`year`、`director` 和 `duration`。

# --hints--

你應該在你的 `__init__` 方法中呼叫 `super().__init__(title, year, director, duration)`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").find_function("__init__").has_stmt("super().__init__(title, year, director, duration)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
class TVSeries(Movie):
    def __init__(self, title, year, director, duration, seasons, total_episodes):
        
        self.seasons = seasons
        self.total_episodes = total_episodes
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
