---
id: 68b80a995417adac579f9545
title: 第 19 步
challengeType: 20
dashedName: step-19
---

# --description--

現在，透過驗證 `seasons` 和 `total_episodes` 來完成你的 `__init__` 方法。創建一個 `if` 述語，檢查 `seasons` 是否小於 1，並引發帶有訊息 `Seasons must be 1 or greater` 的 `ValueError`。

然後，對 `total_episodes` 執行相同的驗證。這次使用訊息 `Total episodes must be 1 or greater`。

# --hints--

當 `seasons` 小於 `1` 時，你應該引發一個帶有訊息 `Seasons must be 1 or greater` 的 `ValueError`。

```js
({
  test: () => runPython(`
  _msg = 'Seasons must be 1 or greater'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 0, 182)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, -1, 182)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

當 `seasons` 是有效數字時，你不應該引發 `ValueError`。

```js
({
    test: () => runPython(`
    try:
      TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 1, 182)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid seasons'
    `)
})
```

當 `total_episodes` 小於 `1` 時，你應該引發一個帶有訊息 `Total episodes must be 1 or greater` 的 `ValueError`。

```js
({
  test: () => runPython(`
  _msg = 'Total episodes must be 1 or greater'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 0)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, -1)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

當 `total_episodes` 是有效數字時，你不應該引發 `ValueError`。

```js
({
    test: () => runPython(`
    try:
      TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 1)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid total_episodes'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class TVSeries(Movie):
    def __init__(self, title, year, director, duration, seasons, total_episodes):
        super().__init__(title, year, director, duration)
--fcc-editable-region--
        

        self.seasons = seasons
        self.total_episodes = total_episodes
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
