---
id: 68b95fa5f6e575879105e74f
title: 步驟 32
challengeType: 20
dashedName: step-32
---

# --description--

既然 `MediaError` 類別已完成，請修改 `add` 方法中的 `if` 述語。將 `TypeError` 取代為使用相同訊息的 `MediaError`。

此外，在 `message` 之後，將 `media_item` 作為第二個引數傳入。這將允許你將導致錯誤的項目儲存在例外物件中。

# --hints--

當 `media_item` 不是 `Movie` 或其子類別的實體時，你應該引發一個帶有訊息 `Only Movie or TVSeries instances can be added` 的 `MediaError`。

```js
({
  test: () => runPython(`
  _msg = 'Only Movie or TVSeries instances can be added'
  _expected_error = f"Expected to raise 'MediaError: {_msg}'"
  _c = MediaCatalogue()
  try:    
    _c.add("Not a movie")
  except MediaError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    _c.add({})
  except MediaError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

你應該將 `media_item` 作為第二個引數傳入 `MediaError`。

```js
({
  test: () => runPython(`
  _c = MediaCatalogue()
  _obj = "Not a movie"
  try:    
    _c.add(_obj)
  except MediaError as e:
    assert e.obj == _obj
  else:
    assert False
  `)
})
```

當 `media_item` 是有效的項目時，你不應該引發 `MediaError`。

```js
({
    test: () => runPython(`
    _c = MediaCatalogue()
    try:
      _c.add(Movie('The Matrix', 1999, 'The Wachowskis', 136))
    except MediaError:
      assert False, 'Expected not to raise a MediaError with valid media_item'
    except TypeError:
      pass

    try:
      _c.add(TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 182))
    except MediaError:
      assert False, 'Expected not to raise a MediaError with valid media_item'
    except TypeError:
      pass
    `)
})
```

你的 `add` 方法不應該引發 `TypeError`。

```js
({
    test: () => runPython(`
    _c = MediaCatalogue()
    try:
      _c.add(Movie('The Matrix', 1999, 'The Wachowskis', 136))
    except TypeError:
      assert False, 'Expected not to raise a TypeError from add method'

    try:
      _c.add(TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 182))
    except TypeError:
      assert False, 'Expected not to raise a TypeError from add method'

    try:
      _c.add("Not a movie")
    except TypeError:
      assert False, 'Expected not to raise a TypeError from add method'
    except MediaError:
      pass
    `)
})
```

# --seed--

## --seed-contents--

```py
class MediaError(Exception):
    """Custom exception for media-related errors."""

    def __init__(self, message, obj):
        super().__init__(message)
        self.obj = obj

class Movie:
    """Parent class representing a movie."""
    
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class TVSeries(Movie):
    """Child class representing an entire TV series."""

    def __init__(self, title, year, director, duration, seasons, total_episodes):
        super().__init__(title, year, director, duration)

        if seasons < 1:
            raise ValueError('Seasons must be 1 or greater')
        if total_episodes < 1:
            raise ValueError('Total episodes must be 1 or greater')
        
        self.seasons = seasons
        self.total_episodes = total_episodes

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.seasons} seasons, {self.total_episodes} episodes, {self.duration} min avg, {self.director}'

class MediaCatalogue:
    """A catalogue that can store different types of media items."""

    def __init__(self):
        self.items = []

    def add(self, media_item):
--fcc-editable-region--
        if not isinstance(media_item, Movie):
            raise TypeError('Only Movie or TVSeries instances can be added')
--fcc-editable-region--
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)

    series1 = TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 182)
    catalogue.add(series1)
    series2 = TVSeries('Breaking Bad', 2008, 'Vince Gilligan', 47, 5, 62)
    catalogue.add(series2)
    
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
