---
id: 68b99289c744bffa96807d04
title: 步驟 37
challengeType: 20
dashedName: step-37
---

# --description--

目前，當你列印目錄時，電影和電視劇會一起顯示。在接下來的步驟中，你將負責將項目分為兩類並分別顯示。

在 `MediaCatalogue` 類別中，創建一個名為 `get_movies` 的方法，並使其傳回一個列表，該列表僅包含可在 `self.items` 中找到的父類別 `Movie` 的實體。

你可以使用你偏好的方法來做到這一點，但請記得使用 `type` 函式，因為 `isinstance` 無法區分父類別和子類別。

# --hints--

你的 `MediaCatalogue` 類別應該有一個名為 `get_movies` 的方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").has_function("get_movies")`)
})
```

你的 `get_movies` 方法應該傳回一個列表，該列表包含 `self.items` 中所有為 `Movie` 類別實體的項目。

```js
({
  test: () => runPython(`
    m1 = Movie("The Godfather", 1972, "Francis Ford Coppola", 175)
    m2 = Movie("Pulp Fiction", 1994, "Quentin Tarantino", 154)
    s1 = TVSeries("Game of Thrones", 2011, "David Benioff & D.B. Weiss", 57, 8, 73)
    s2 = TVSeries("Stranger Things", 2016, "The Duffer Brothers", 50, 4, 34)
    c = MediaCatalogue()
    c.add(m1)
    c.add(s1)
    c.add(m2)
    c.add(s2)
    movies = c.get_movies()
    assert len(movies) == 2
    assert m1 in movies and m2 in movies
  `)
})
```

# --seed--

## --seed-contents--

```py
class MediaError(Exception):
    """Custom exception for media-related errors."""

    def __init__(self, message, obj):
        super().__init__(message)
        self.obj = obj

class Movie:
    """Parent class representing a movie."""
    
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class TVSeries(Movie):
    """Child class representing an entire TV series."""

    def __init__(self, title, year, director, duration, seasons, total_episodes):
        super().__init__(title, year, director, duration)

        if seasons < 1:
            raise ValueError('Seasons must be 1 or greater')
        if total_episodes < 1:
            raise ValueError('Total episodes must be 1 or greater')
        
        self.seasons = seasons
        self.total_episodes = total_episodes

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.seasons} seasons, {self.total_episodes} episodes, {self.duration} min avg, {self.director}'

class MediaCatalogue:
    """A catalogue that can store different types of media items."""

    def __init__(self):
        self.items = []

    def add(self, media_item):
        if not isinstance(media_item, Movie):
            raise MediaError('Only Movie or TVSeries instances can be added', media_item)
        self.items.append(media_item)
--fcc-editable-region--
    
--fcc-editable-region--
    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)

    series1 = TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 182)
    catalogue.add(series1)
    series2 = TVSeries('Breaking Bad', 2008, 'Vince Gilligan', 47, 5, 62)
    catalogue.add(series2)

    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
except MediaError as e:
    print(f'Media Error: {e}')
    print(f'Unable to add {e.obj}: {type(e.obj)}')

```
