---
id: 68b95956da014b1a8e32ffb3
title: 第 30 步
challengeType: 20
dashedName: step-30
---

# --description--

在先前的模組中，你學到你可以透過繼承內建的 `Exception` 類別來建置你自己的訂製例外：

```py
class CustomError(Exception):
    ...
```

創建一個名為 `MediaError` 的類別。讓它繼承自 `Exception` 類別，並為它添加一個說明字串，內容為 `Custom exception for media-related errors.`。

# --hints--

你應該創建一個名為 `MediaError` 的類別。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("MediaError")`)
})
```

你的 `MediaError` 類別應該繼承自 `Exception`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaError").inherits_from("Exception")`)
})
```

你應該為你的 `MediaError` 類別添加 docstring `Custom exception for media-related errors.`。

```js
({
    test: () => runPython(`
    assert MediaError.__doc__ == 'Custom exception for media-related errors.'
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Movie:
    """Parent class representing a movie."""
    
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class TVSeries(Movie):
    """Child class representing an entire TV series."""

    def __init__(self, title, year, director, duration, seasons, total_episodes):
        super().__init__(title, year, director, duration)

        if seasons < 1:
            raise ValueError('Seasons must be 1 or greater')
        if total_episodes < 1:
            raise ValueError('Total episodes must be 1 or greater')
        
        self.seasons = seasons
        self.total_episodes = total_episodes

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.seasons} seasons, {self.total_episodes} episodes, {self.duration} min avg, {self.director}'

class MediaCatalogue:
    """A catalogue that can store different types of media items."""

    def __init__(self):
        self.items = []

    def add(self, media_item):
        if not isinstance(media_item, Movie):
            raise TypeError('Only Movie or TVSeries instances can be added')
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)

    series1 = TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 182)
    catalogue.add(series1)
    series2 = TVSeries('Breaking Bad', 2008, 'Vince Gilligan', 47, 5, 62)
    catalogue.add(series2)

    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
