---
id: 6a89c69c8dbfda74151fa0b6
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 워크숍에서는 재귀를 사용해 `1`부터 주어진 숫자까지의 숫자 목록을 만드는 `countup` 함수를 만드세요.

`countup`이라는 이름의 함수를 정의하고 `number`라는 매개변수를 하나 추가하는 것부터 시작하세요. 함수 본문에 `pass` 키워드를 넣어 함수가 비어 있어도 유효하도록 만드세요.

# --hints--

`countup`이라는 이름의 함수가 있고 `number`라는 매개변수를 포함해야 합니다.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("countup")
assert not _func.is_empty()
assert _func.has_args("number")
`) })
```

`countup` 함수에는 `pass` 키워드가 포함되어야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_pass()`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
