---
id: 6a89c69c99298438d6fbe60b
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

이제 재귀 경우를 추가하세요. `if` 문 다음에 `count_list`라는 변수를 만드세요. `number - 1`을 인수로 `countup`을 호출한 결과를 할당하세요.

각 호출은 더 작은 숫자를 사용하므로 재귀가 기본 경우로 진행됩니다.

# --hints--

`countup(number - 1)`을 `count_list`라는 변수에 할당하세요.

```js
({ test: () => runPython(`
_variable = _Node(_code).find_function("countup").find_variable("count_list")
assert _variable.is_equivalent("count_list = countup(number - 1)")
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
--fcc-editable-region--
    
--fcc-editable-region--
```
