---
id: 6a89c69c6af682856558d78b
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

재귀 호출은 더 작은 숫자를 포함하는 리스트를 반환합니다. `append()`를 사용해 `number`를 `count_list`에 추가한 다음, `count_list`를 반환하세요.

# --hints--

`number`를 `count_list`에 추가해야 합니다.

```js
({ test: () => runPython(`
_calls = _Node(_code).find_function("countup").find_calls("append")
assert len(_calls) == 1
assert _calls[0].is_equivalent("count_list.append(number)")
`) })
```

`countup` 함수에서 `count_list`를 반환해야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_return("count_list")`) })
```

`countup` 함수는 `3`을 인자로 호출했을 때 `[1, 2, 3]`을 반환해야 합니다.

```js
({ test: () => runPython(`assert countup(3) == [1, 2, 3]`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
--fcc-editable-region--
    
--fcc-editable-region--
```
