---
id: 6a89c69c111694edc6d1c628
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

`print()` 호출 안에 `5`를 전달하여 `countup` 함수를 테스트하세요. 출력은 `[1, 2, 3, 4, 5]`여야 합니다.

이로써 countup 함수 워크숍이 완료되었습니다.

# --hints--

`print()` 함수 안에서 `countup(5)`를 호출해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_call("print(countup(5))")
assert countup(5) == [1, 2, 3, 4, 5]
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

print(countup(5))
```
