---
id: 6a89c69c91bc7365fbd4b733
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

`countup` 함수에 기본 사례를 추가하세요. `pass`를 `if` 문으로 바꾸고, `number`가 `1`보다 작은지 확인하세요. `if` 문 안에서는 빈 리스트를 반환하세요.

# --hints--

`number`가 `1`보다 작을 때 `countup` 함수는 빈 리스트를 반환해야 합니다.

```js
({ test: () => runPython(`assert countup(0) == []`) })
```

`number`가 `1` 이상일 때 `countup` 함수가 빈 리스트를 반환하면 안 됩니다.

```js
({ test: () => runPython(`assert countup(1) is None`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
