---
id: a105e963526e7de52b219be9
title: 實作唯一排序聯集（數學）
challengeType: 26
dashedName: implement-a-unique-sorted-union
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個名為 `uniteUnique` 的函式（程式）。
1. `uniteUnique` 函式（程式）應該接受兩個或以上的陣列作為引數。
1. 該函式應傳回一個新的陣列，該陣列包含引數陣列中首次出現的唯一值，且順序依其在引數中首次出現的順序排列。 例如，輸入像 `[1, 2, 4], [2, 3, 5]` 的話，輸出將會是 `[1, 2, 4, 3, 5]`。

# --hints--

你應該有一個 `uniteUnique` 函式（程式）。

```js
assert.isFunction(uniteUnique);
```

`uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1])` 應該傳回 `[1, 3, 2, 5, 4]`。

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

`uniteUnique([1, 2, 3], [5, 2, 1])` 應該傳回 `[1, 2, 3, 5]`。

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1]), [1, 2, 3, 5]);
```

`uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8])` 應該傳回 `[1, 2, 3, 5, 4, 6, 7, 8]`。

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8]), [
  1,
  2,
  3,
  5,
  4,
  6,
  7,
  8
]);
```

`uniteUnique([1, 3, 2], [5, 4], [5, 6])` 應該傳回 `[1, 3, 2, 5, 4, 6]`。

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 4], [5, 6]), [1, 3, 2, 5, 4, 6]);
```

`uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1])` 應該傳回 `[1, 3, 2, 5, 4]`。

```js
assert.deepEqual(uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const uniteUnique = (...arrays) => {
  return arrays.reduce((a, b) => {
    return [
      ...a,
      ...b.filter((e, currentIndex) => b.indexOf(e) === currentIndex && a.indexOf(e) === -1)
    ];
  }, []);
};
```
