---
id: 587d78b1367417b2b2512b0c
title: レスポンシブなタイポグラフィを作る
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/crzN7T8'
forumTopicId: 301141
dashedName: make-typography-responsive
---

# --description--

`em` や `px` を使ってテキストのサイズを指定する代わりに、ビューポート単位を使ってレスポンシブなタイポグラフィを作ることができます。ビューポート単位はパーセンテージのように相対単位ですが、基準が異なります。ビューポート単位はデバイスのビューポートの幅や高さに対して相対的で、パーセンテージは親コンテナ要素のサイズに対して相対的です。

4つの異なるビューポート単位は次の通りです：

<ul><li><code>vw</code>（ビューポート幅）：<code>10vw</code> はビューポートの幅の10％です。</li><li><code>vh</code>（ビューポート高さ）：<code>3vh</code> はビューポートの高さの3％です。</li><li><code>vmin</code>（ビューポート最小値）：<code>70vmin</code> はビューポートの小さい方の寸法（高さまたは幅）の70％です。</li><li><code>vmax</code>（ビューポート最大値）：<code>100vmax</code> はビューポートの大きい方の寸法（高さまたは幅）の100％です。</li></ul>

以下は `body` タグのサイズをビューポート幅の30％に設定する例です。

```css
body {
  width: 30vw;
}
```

# --instructions--

`width` タグの `h2` をビューポート幅の80％に、段落の `width` をビューポートの小さい方の寸法の75％に設定してください。

# --hints--

`h2` タグの `width` は 80vw にしてください。

```js
assert.match(
  __helpers.removeCssComments(code),
  /h2\s*?{\s*?width:\s*?80vw;\s*?}/g
);
```

`p` タグの `width` は 75vmin にしてください。

```js
assert.match(
  __helpers.removeCssComments(code),
  /p\s*?{\s*?width:\s*?75vmin;\s*?}/g
);
```

# --seed--

## --seed-contents--

```html
<style>

</style>

<h2>Importantus Ipsum</h2>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```

# --solutions--

```html
<style>
  h2 {
    width: 80vw;
  }
  p {
    width: 75vmin;
  }
</style>

<h2>Importantus Ipsum</h2>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```
