---
id: 587d78b1367417b2b2512b09
title: 画像をレスポンシブにする
challengeType: 0
forumTopicId: 301140
dashedName: make-an-image-responsive
---

# --description--

CSSで画像をレスポンシブにするのは実はとても簡単です。画像に次のプロパティを追加するだけです。

```css
img {
  max-width: 100%;
  height: auto;
}
```

`max-width`の`100%`は、画像が所属するコンテナよりも幅が広くならないようにし、`height`の`auto`は画像の元のアスペクト比を維持します。

# --instructions--

`responsive-img`クラスにスタイルルールを追加してレスポンシブにしてください。コンテナ（この場合はプレビューウィンドウ）よりも幅が広くならず、元のアスペクト比を保つようにします。コードを追加したら、プレビューのサイズを変更して画像の動作を確認してください。

# --hints--

`responsive-img`クラスには`max-width`を`100%`に設定してください。

```js
assert.strictEqual(getComputedStyle(document.querySelector('.responsive-img')).maxWidth, '100%');
```

`responsive-img`クラスには`height`を`auto`に設定してください。

```js
assert.match(code, /height:\s*?auto;/g);
```

# --seed--

## --seed-contents--

```html
<style>
.responsive-img {


}

img {
  width: 600px;
}
</style>

<img
  class="responsive-img"
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
```

# --solutions--

```html
<style>
  .responsive-img {
    max-width: 100%;
    height: auto;
  }

  img {
    width: 600px;
  }
</style>

<img
  class="responsive-img"
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
```
