---
id: 587d78b1367417b2b2512b0a
title: 高解像度ディスプレイ用にレティナ画像を使う
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/cVZ4Rfp'
forumTopicId: 301142
dashedName: use-a-retina-image-for-higher-resolution-displays
---

# --description--

インターネットに接続されたデバイスが増えるにつれて、そのサイズや仕様はさまざまで、使用されるディスプレイも外観や内部構造が異なる場合があります。ピクセル密度はデバイスごとに異なることがあり、この密度はピクセル毎インチ（PPI）またはドット毎インチ（DPI）として知られています。最も有名なディスプレイの一つが、最新のApple MacBook Proノートパソコンや最近のiMacコンピューターで使われている「レティナディスプレイ」です。「レティナ」と「非レティナ」ディスプレイ間のピクセル密度の違いにより、高解像度ディスプレイを想定して作られていない画像は、高解像度ディスプレイ上で「ピクセル化」して見えることがあります。

MacBook Proの「レティナディスプレイ」などの高解像度ディスプレイで画像を適切に表示させる最も簡単な方法は、`width` と `height` の値を元のファイルの半分に設定することです。以下は、元の高さと幅の半分だけを使っている画像の例です。

```html
<style>
  img {
    height: 250px;
    width: 250px;
  }
</style>
<img src="coolPic500x500" alt="A most excellent picture" />
```

# --instructions--

`width` タグの `height` と `img` を元の値の半分に設定してください。この場合、元の `height` と元の `width` はどちらも `200px` です。

# --hints--

`img` タグの `width` は100ピクセルにしてください。

```js
assert.strictEqual(document.querySelector('img').width, 100);
```

`img` タグの `height` は100ピクセルにしてください。

```js
assert.strictEqual(document.querySelector('img').height, 100);
```

# --seed--

## --seed-contents--

```html
<style></style>

<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickers-CamperBot200x200.jpg"
  alt="freeCodeCamp sticker that says 'Because CamperBot Cares'"
/>
```

# --solutions--

```html
<style>
  img {
    height: 100px;
    width: 100px;
  }
</style>

<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickers-CamperBot200x200.jpg"
  alt="freeCodeCamp sticker that says 'Because CamperBot Cares'"
/>
```
