---
id: 587d78b0367417b2b2512b08
title: メディアクエリを作成する
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/cqwKrtm'
forumTopicId: 301139
dashedName: create-a-media-query
---

# --description--

メディアクエリはCSS3で導入された新しい手法で、異なるビューポートサイズに基づいてコンテンツの表示を変更します。ビューポートとはユーザーがウェブページで見ている表示エリアのことで、アクセスするデバイスによって異なります。

メディアクエリはメディアタイプで構成され、そのメディアタイプがドキュメントが表示されているデバイスのタイプと一致すると、スタイルが適用されます。メディアクエリの中には好きなだけセレクターやスタイルを入れることができます。

以下は、デバイスの幅が`100px`以下の場合にコンテンツを返すメディアクエリの例です。

```css
@media (max-width: 100px) {
  /* CSS Rules */
}
```

次のメディアクエリは、デバイスの高さが`350px`以上の場合にコンテンツを返します。

```css
@media (min-height: 350px) {
  /* CSS Rules */
}
```

メディアクエリ内のCSSは、メディアタイプが使用中のデバイスのタイプと一致した場合にのみ適用されることを覚えておいてください。

# --instructions--

メディアクエリを追加して、デバイスの高さが`p`以下の場合に`font-size`タグの`10px`を`800px`にしてください。

# --hints--

デバイスの`@media`が`height`以下の場合の`800px`クエリを宣言する必要があります。

```js
const media = new __helpers.CSSHelp(document).getCSSRules('media');
assert.isTrue(
  media.some(x => x.media?.mediaText?.includes('(max-height: 800px)'))
);
```

デバイスの`p`が`font-size`以下の場合、`10px`要素の`height`を`800px`にしてください。

```js
const rules = new __helpers.CSSHelp(document).getRuleListsWithinMedia(
  '(max-height: 800px)'
);
assert.strictEqual(
  rules?.find(x => x.selectorText === 'p')?.style?.fontSize,
  '10px'
);
```

デバイスの`p`が`font-size`より大きい場合、`20px`要素の初期`height`を`800px`にしてください。

```js
const ifPFirst = new __helpers.CSSHelp(document)
  .getCSSRules()
  ?.find(x => x?.selectorText === 'p' || x?.media);
assert.strictEqual(ifPFirst?.style?.fontSize, '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  p {
    font-size: 20px;
  }

  /* Only change code below this line */

  /* Only change code above this line */
</style>

<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```

# --solutions--

```html
<style>
  p {
    font-size: 20px;
  }

  @media (max-height: 800px) {
    p {
      font-size: 10px;
    }
  }
</style>

<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```
