---
id: 587d7b84367417b2b2512b35
title: Individuare nomi di variabili e funzioni scritti male
challengeType: 1
forumTopicId: 301186
dashedName: catch-misspelled-variable-and-function-names
---

# --description--

I metodi `console.log()` e `typeof` sono i due modi principali per controllare valori intermedi e tipi di output di un programma. Ora è il momento di affrontare le forme comuni che assumono i bug. Un problema a livello di sintassi con cui chi digita velocemente può empatizzare è il semplice errore di ortografia.

Caratteri trasposti, mancanti o con maiuscole/minuscole errate in un nome di variabile o funzione faranno sì che il browser cerchi un oggetto che non esiste — e segnalerà un errore di riferimento. I nomi di variabili e funzioni in JavaScript sono sensibili alle maiuscole.

# --instructions--

Correggi i due errori di ortografia nel codice affinché il calcolo `netWorkingCapital` funzioni.

# --hints--

Controlla l’ortografia delle due variabili usate nel calcolo netWorkingCapital, l’output della console dovrebbe mostrare che "Net working capital is: 2".

```js
assert(netWorkingCapital === 2);
```

Non dovrebbero esserci variabili scritte male nel codice.

```js
assert(!__helpers.removeJSComments(code).match(/recievables/g));
```

La variabile `receivables` dovrebbe essere dichiarata e usata correttamente nel codice.

```js
assert(__helpers.removeJSComments(code).match(/receivables/g).length == 2);
```

Non dovrebbero esserci variabili scritte male nel codice.

```js
assert(!__helpers.removeJSComments(code).match(/payable;/g));
```

La variabile `payables` dovrebbe essere dichiarata e usata correttamente nel codice.

```js
assert(__helpers.removeJSComments(code).match(/payables/g).length == 2);
```

# --seed--

## --seed-contents--

```js
let receivables = 10;
let payables = 8;
let netWorkingCapital = recievables - payable;
console.log(`Net working capital is: ${netWorkingCapital}`);
```

# --solutions--

```js
let receivables = 10;
let payables = 8;
let netWorkingCapital = receivables - payables;
console.log(`Net working capital is: ${netWorkingCapital}`);
```
