---
id: 587d7b85367417b2b2512b39
title: Rilevare la mancanza di parentesi tonde aperte e chiuse dopo una chiamata a funzione
challengeType: 1
forumTopicId: 301185
dashedName: catch-missing-open-and-closing-parenthesis-after-a-function-call
---

# --description--

Quando una funzione o un metodo non prende argomenti, potresti dimenticare di includere le parentesi tonde aperte e chiuse (vuote) quando la chiami. Spesso il risultato di una chiamata a funzione viene salvato in una variabile per un uso successivo nel tuo codice. Questo errore può essere individuato registrando i valori delle variabili (o i loro tipi) nella console e notando che una è impostata su un riferimento a funzione, invece del valore previsto che la funzione restituisce.

Le variabili nell'esempio seguente sono diverse:

```js
function myFunction() {
  return "You rock!";
}
let varOne = myFunction;
let varTwo = myFunction();
```

Qui `varOne` è la funzione `myFunction`, e `varTwo` è la stringa `You rock!`.

# --instructions--

Correggi il codice in modo che la variabile `result` sia impostata sul valore restituito dalla chiamata della funzione `getNine`.

# --hints--

Il tuo codice dovrebbe correggere la variabile `result` in modo che sia impostata sul numero che la funzione `getNine` restituisce.

```js
assert(result == 9);
```

Il tuo codice dovrebbe chiamare la funzione `getNine`.

```js
assert(__helpers.removeJSComments(code).match(/getNine\(\)/g).length == 2);
```

# --seed--

## --seed-contents--

```js
function getNine() {
  let x = 6;
  let y = 3;
  return x + y;
}

let result = getNine;
console.log(result);
```

# --solutions--

```js
function getNine() {
 let x = 6;
 let y = 3;
 return x + y;
}

let result = getNine();
console.log(result);
```
