---
id: 587d7b84367417b2b2512b37
title: "Gestire l'uso misto di virgolette singole e doppie"
challengeType: 1
forumTopicId: 301188
dashedName: catch-mixed-usage-of-single-and-double-quotes
---

# --description--

JavaScript permette di usare sia virgolette singole (`'`) sia virgolette doppie (`"`) per dichiarare una stringa. La scelta di quale usare dipende generalmente dalla preferenza personale, con alcune eccezioni.

Avere due opzioni è utile quando una stringa contiene contrazioni o un altro testo tra virgolette. Basta fare attenzione a non chiudere la stringa troppo presto, perché questo causa un errore di sintassi.

Ecco alcuni esempi di uso misto delle virgolette:

```js
const grouchoContraction = "I've had a perfectly wonderful evening, but this wasn't it.";
const quoteInString = "Groucho Marx once said 'Quote me as saying I was mis-quoted.'";
const uhOhGroucho = 'I've had a perfectly wonderful evening, but this wasn't it.';
```

I primi due sono corretti, ma il terzo è errato.

Naturalmente, va bene usare solo uno stile di virgolette. Puoi fare escape delle virgolette all’interno della stringa usando il carattere di escape barra rovesciata (`\`):

```js
const allSameQuotes = 'I\'ve had a perfectly wonderful evening, but this wasn\'t it.';
```

# --instructions--

Correggi la stringa in modo che usi virgolette diverse per il valore `href` oppure fai escape di quelle esistenti. Mantieni le virgolette doppie intorno all’intera stringa.

# --hints--

Il tuo codice dovrebbe correggere le virgolette intorno al valore `href` `#Home` cambiandole o facendo escape.

```js
assert(__helpers.removeJSComments(code).match(/<a href=\s*?('|\\")#Home\1\s*?>/g));
```

Il tuo codice dovrebbe mantenere le virgolette doppie intorno all’intera stringa.

```js
assert(__helpers.removeJSComments(code).match(/"<p>.*?<\/p>";/g));
```

# --seed--

## --seed-contents--

```js
let innerHtml = "<p>Click here to <a href="#Home">return home</a></p>";
console.log(innerHtml);
```

# --solutions--

```js
let innerHtml = "<p>Click here to <a href=\"#Home\">return home</a></p>";
console.log(innerHtml);
```
