---
id: 587d7b84367417b2b2512b36
title: 'Individua parentesi, parentesi quadre, parentesi graffe e virgolette non chiuse'
challengeType: 1
forumTopicId: 301190
dashedName: catch-unclosed-parentheses-brackets-braces-and-quotes
---

# --description--

Un altro errore di sintassi da tenere presente è che tutte le parentesi tonde, parentesi quadre, parentesi graffe e virgolette aperte devono avere la loro corrispondente chiusura. Dimenticare una parte succede spesso quando modifichi codice esistente e inserisci elementi con uno dei tipi di coppia. Inoltre, fai attenzione quando annidi blocchi di codice dentro altri, come aggiungere una funzione di callback come argomento a un metodo.

Un modo per evitare questo errore è, non appena digiti il carattere di apertura, includere subito la corrispondente chiusura, poi spostare il cursore tra i due e continuare a scrivere il codice. Fortunatamente, la maggior parte degli editor di codice moderni genera automaticamente la seconda metà della coppia.

# --instructions--

Correggi i due errori di coppia nel codice.

# --hints--

Il tuo codice dovrebbe correggere la parte mancante dell'array.

```js
assert(__helpers.removeJSComments(code).match(/myArray\s*?=\s*?\[\s*?1\s*?,\s*?2\s*?,\s*?3\s*?\];/g));
```

Il tuo codice dovrebbe correggere la parte mancante del metodo `.reduce()`. L'output della console dovrebbe mostrare che `Sum of array values is: 6`.

```js
assert(arraySum === 6);
```

# --seed--

## --seed-contents--

```js
let myArray = [1, 2, 3;
let arraySum = myArray.reduce((previous, current =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```

# --solutions--

```js
let myArray = [1, 2, 3];
let arraySum = myArray.reduce((previous, current) =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```
