---
id: 587d7b86367417b2b2512b3b
title: "Rilevare errori off by one quando si usa l'indicizzazione"
challengeType: 1
forumTopicId: 301189
dashedName: catch-off-by-one-errors-when-using-indexing
---

# --description--

<dfn>Errori off by one</dfn> (a volte chiamati OBOE) si verificano quando cerchi di consultare un indice specifico di una stringa o di un array (per estrarre o consultare un segmento), o quando iteri sugli indici di questi. L'indicizzazione in JavaScript parte da zero, non da uno, il che significa che l'ultimo indice è sempre uno in meno della lunghezza dell'elemento. Se provi a consultare un indice uguale alla lunghezza, il programma potrebbe lanciare un errore di riferimento "indice fuori intervallo" o stampare `undefined`.

Quando usi metodi di stringa o array che prendono intervalli di indici come argomenti, è utile leggere la documentazione e capire se sono inclusivi (l'elemento all'indice dato fa parte di ciò che viene restituito) o no. Ecco alcuni esempi di errori off by one:

```js
let alphabet = "abcdefghijklmnopqrstuvwxyz";
let len = alphabet.length;
for (let i = 0; i <= len; i++) {
  console.log(alphabet[i]);
}
for (let j = 1; j < len; j++) {
  console.log(alphabet[j]);
}
for (let k = 0; k < len; k++) {
  console.log(alphabet[k]);
}
```

Il primo esempio qui itera una volta di troppo, e il secondo itera una volta di meno (mancando il primo indice, 0). Il terzo esempio è corretto.

# --instructions--

Correggi i due errori di indicizzazione nella funzione seguente in modo che tutti i numeri da 1 a 5 vengano stampati nella console.

# --hints--

Il tuo codice dovrebbe impostare la condizione iniziale del loop in modo che inizi dal primo indice.

```js
assert(__helpers.removeJSComments(code).match(/i\s*?=\s*?0\s*?;/g).length == 1);
```

Il tuo codice dovrebbe correggere la condizione iniziale del loop in modo che l'indice inizi da 0.

```js
assert(!__helpers.removeJSComments(code).match(/i\s?=\s*?1\s*?;/g));
```

Il tuo codice dovrebbe impostare la condizione terminale del loop in modo che si fermi all'ultimo indice.

```js
assert(__helpers.removeJSComments(code).match(/i\s*<\s*len\s*;|i\s*<=\s*len\s*-\s*1\s*;/g).length == 1);
```

Il tuo codice dovrebbe correggere la condizione terminale del loop in modo che si fermi 1 prima della lunghezza.

```js
assert(!__helpers.removeJSComments(code).match(/i\s*?<=\s*?len;/g));
```

# --seed--

## --seed-contents--

```js
function countToFive() {
  let firstFive = "12345";
  let len = firstFive.length;
  // Only change code below this line
  for (let i = 1; i <= len; i++) {
  // Only change code above this line
    console.log(firstFive[i]);
  }
}

countToFive();
```

# --solutions--

```js
function countToFive() {
 let firstFive = "12345";
 let len = firstFive.length;
 // Only change code below this line
 for (let i = 0; i < len; i++) {
 // Only change code above this line
   console.log(firstFive[i]);
 }
}

countToFive();
```
