---
id: 587d7b85367417b2b2512b3a
title: "Rilevare argomenti passati nell'ordine sbagliato quando si chiama una funzione"
challengeType: 1
forumTopicId: 301184
dashedName: catch-arguments-passed-in-the-wrong-order-when-calling-a-function
---

# --description--

Continuando la discussione sulle chiamate di funzione, il prossimo bug a cui fare attenzione è quando gli argomenti di una funzione vengono forniti nell'ordine sbagliato. Se gli argomenti sono di tipi diversi, come una funzione che si aspetta un array e un numero intero, questo probabilmente genererà un errore di runtime. Se gli argomenti sono dello stesso tipo (tutti numeri interi, per esempio), allora la logica del codice non avrà senso. Assicurati di fornire tutti gli argomenti richiesti, nell'ordine corretto per evitare questi problemi.

# --instructions--

La funzione `raiseToPower` eleva una base a un esponente. Sfortunatamente, non viene chiamata correttamente: correggi il codice in modo che il valore di `power` sia l'8 previsto.

# --hints--

Il tuo codice dovrebbe correggere la variabile `power` in modo che sia uguale a 2 elevato alla terza potenza, non 3 elevato alla seconda.

```js
assert(power == 8);
```

Il tuo codice dovrebbe usare l'ordine corretto degli argomenti per la chiamata della funzione `raiseToPower`.

```js
assert(__helpers.removeJSComments(code).match(/raiseToPower\(\s*?base\s*?,\s*?exp\s*?\);/g));
```

# --seed--

## --seed-contents--

```js
function raiseToPower(b, e) {
  return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(exp, base);
console.log(power);
```

# --solutions--

```js
function raiseToPower(b, e) {
 return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(base, exp);
console.log(power);
```
