---
id: 587d7b83367417b2b2512b33
title: Usa la console JavaScript per controllare il valore di una variabile
challengeType: 1
forumTopicId: 18372
dashedName: use-the-javascript-console-to-check-the-value-of-a-variable
---

# --description--

Sia Chrome che Firefox hanno ottime console JavaScript, conosciute anche come strumenti per sviluppatori, per eseguire il debug del tuo JavaScript.

Puoi trovare gli strumenti per sviluppatori nel menu di Chrome o la console web nel menu di Firefox. Se usi un browser diverso o un telefono cellulare, ti consigliamo vivamente di passare a Firefox o Chrome su desktop.

Il metodo `console.log()`, che "stampa" l’output di ciò che si trova tra le sue parentesi tonde nella console, sarà probabilmente lo strumento di debugging più utile. Posizionarlo in punti strategici del tuo codice può mostrarti i valori intermedi delle variabili. È buona pratica avere un’idea di quale dovrebbe essere l’output prima di guardare qual è realmente. Avere punti di controllo per vedere lo stato dei tuoi calcoli durante tutto il codice ti aiuterà a restringere dove si trova il problema.

Ecco un esempio per stampare la stringa `Hello world!` nella console:

```js
console.log('Hello world!');
```

# --instructions--

Usa il metodo `console.log()` per stampare il valore della variabile `a` dove indicato nel codice.

# --hints--

Il tuo codice dovrebbe usare `console.log()` per controllare il valore della variabile `a`.

```js
assert(__helpers.removeJSComments(code).match(/console\.log\(a\)/g));
```

# --seed--

## --seed-contents--

```js
let a = 5;
let b = 1;
a++;
// Only change code below this line


let sumAB = a + b;
console.log(sumAB);
```

# --solutions--

```js
var a = 5; console.log(a);
```
