---
id: 587d7b86367417b2b2512b3c
title: Usa cautela quando reimposti variabili dentro un loop
challengeType: 1
forumTopicId: 301194
dashedName: use-caution-when-reinitializing-variables-inside-a-loop
---

# --description--

A volte è necessario salvare informazioni, incrementare contatori o reimpostare variabili all'interno di un loop. Un problema potenziale si verifica quando le variabili dovrebbero essere reimpostate e non lo sono, o viceversa. Questo è particolarmente pericoloso se reimposti accidentalmente la variabile usata per la condizione terminale, causando un loop infinito.

Stampare i valori delle variabili a ogni ciclo del loop usando `console.log()` può far emergere comportamenti con bug legati al reimpostare o al non reimpostare una variabile.

# --instructions--

La funzione seguente dovrebbe creare un array bidimensionale con `m` righe e `n` colonne di zeri. Sfortunatamente, non produce l'output previsto perché la variabile `row` non viene reimpostata (riportata a un array vuoto) nel loop esterno. Correggi il codice in modo che ritorni un array 3x2 corretto di zeri, che appare come `[[0, 0], [0, 0], [0, 0]]`.

# --hints--

Il tuo codice dovrebbe impostare la variabile `matrix` su un array che contiene 3 righe di 2 colonne di zeri ciascuna.

```js
assert(JSON.stringify(matrix) == '[[0,0],[0,0],[0,0]]');
```

La variabile `matrix` dovrebbe avere 3 righe.

```js
assert(matrix.length == 3);
```

La variabile `matrix` dovrebbe avere 2 colonne in ogni riga.

```js
assert(
  matrix[0].length == 2 && matrix[1].length === 2 && matrix[2].length === 2
);
```

`zeroArray(4,3)` dovrebbe restituire un array con 4 righe di 3 colonne di zeri ciascuna.

```js
assert(JSON.stringify(zeroArray(4,3)) == '[[0,0,0],[0,0,0],[0,0,0],[0,0,0]]');
```

# --seed--

## --seed-contents--

```js
function zeroArray(m, n) {
  // Creates a 2-D array with m rows and n columns of zeroes
  let newArray = [];
  let row = [];
  for (let i = 0; i < m; i++) {
    // Adds the m-th row into newArray

    for (let j = 0; j < n; j++) {
      // Pushes n zeroes into the current row to create the columns
      row.push(0);
    }
    // Pushes the current row, which now has n zeroes in it, to the array
    newArray.push(row);
  }
  return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```

# --solutions--

```js
function zeroArray(m, n) {
 // Creates a 2-D array with m rows and n columns of zeroes
 let newArray = [];
 for (let i = 0; i < m; i++) {
   let row = [];
   // Adds the m-th row into newArray

   for (let j = 0; j < n; j++) {
     // Pushes n zeroes into the current row to create the columns
     row.push(0);
   }
   // Pushes the current row, which now has n zeroes in it, to the array
   newArray.push(row);
 }
 return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```
