---
id: 61695ab9f6ffe951c16d03dd
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

Nest a `head` element within the `html` element. Just after the `head` element, add a `body` element.

# --hints--

يجب أن يكون لديك وسم فتح للعنصر `head`.

```js
assert(code.match(/<head\s*>/i));
```

يجب أن يكون لديك وسم أغلاق للعنصر `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

يجب أن يكون لديك وسم فتح للعنصر `body`.

```js
assert(code.match(/<body\s*>/i));
```

يجب أن يكون لديك وسم أغلاق للعنصر `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

يجب أن يكون هناك علامة فتح `head` واحدة فقط.

```js
assert(code.match(/<head\s*>/ig).length === 1);
```

يجب أن يكون هناك علامة أغلاق `head` واحدة فقط.

```js
assert(code.match(/<\/head\s*>/ig).length === 1);
```

يجب أن يكون هناك علامة فتح `body` واحدة فقط.

```js
assert(code.match(/<body\s*>/ig).length === 1);
```

يجب أن يكون هناك علامة أغلاق `body` واحدة فقط.

```js
assert(code.match(/<\/body\s*>/ig).length === 1);
```

يجب أن يكون عنصر `head` فارغاً.

```js
assert(code.match(/<head\s*>\s*<\/head\s*>/i));
```

يجب أن يكون عنصر `body` فارغاً.

```js
assert(code.match(/<body\s*>\s*<\/body\s*>/i));
```

يجب أن يأتي عنصر `body` الخاص بك بعد عنصر `head` الخاص بك.

```js
assert(code.match(/<\/head\s*>\s*<body\s*>/i));
```

يجب أن يكون عنصر `head` داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

يجب أن يكون عنصر `body` داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
