---
id: 61695197ac34f0407e339882
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

As you've seen in the previous projects, webpages should start with a `DOCTYPE html` declaration, followed by an `html` element.

قم بإضافة تعريف `DOCTYPE html` في الجزء العلوي من الوثيقة، وعنصر `html` بعد ذلك. اعطي عنصر `html` سمة `lang` بقيمة `en`.

# --hints--

يجب أن يحتوي الكود الخاص بك على مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

عليك أن تضيف مسافة بعد مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

عليك أن تعرّف أن نوع مستندك هو `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

ينبغي إغلاق إعلان `DOCTYPE` بواسطة <code dir="ltr">&gt;</code> بعد تعيين النوع.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

عنصر `html` الخاص بك يجب أن يكون له opening tag معه سمة `lang` تساوي `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصرك `html` على علامة أغلاق (closing tag).

```js
assert(code.match(/<\/html\s*>/gi));
```

تعريف `DOCTYPE` الخاص بك يجب أن يكون عند بداية HTML الخاص بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
