---
id: 61764c602bee6974e7790f35
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

To give the markers different colors, you will need to add a unique class to each one. Multiple classes can be added to an element by listing them in the `class` attribute and separating them with a space. For example, the following adds both the `animal` and `dog` classes to a `div` element.

```html
<div class="animal dog">
```

إذا قمت بإضافة فئات متعددة إلى عنصر HTML، فإن أنماط الفئات الأولى التي تضعها قد يتم تجاوزها بواسطة الفئات اللاحقة.

أولاً، أضف class يسمي `one` إلى أول العنصر `div` هو marker.


# --hints--

يجب عليك إضافة class قيمته `one` إلى العنصر `div` لأول marker.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

يجب أن يحتوي أول `div` marker على الفئات `marker` و `one`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
