---
id: 617bd6ec666b1da2587e4e37
title: الخطوة 34
challengeType: 0
dashedName: step-34
---

# --description--

There are three more tertiary colors: chartreuse green (green + yellow), azure (blue + cyan), and rose (red + magenta).

لإنشاء chartreuse green، قم بتحديث وظيفة `rgb` في قاعدة `.one` بحيث يكون الأحمر بقيمة `127`، و قيمة الأخضر بالحد الأقصى.

لإنشاء azure، قم بتحديث وظيفة `rgb` في قاعدة `.two` بحيث يكون الاخضر بقيمة `127`، و قيمة الأزرق بالحد الأقصى.

ولإنشاء rose، قم بتحديث وظيفة `rgb` في قاعدة `.three` بحيث يكون الأزرق بقيمة `127`، و قيمة الأحمر بالحد الأقصى.

# --hints--

يجب أن يكون لقاعدة CSS لـ `.one` الخاصية `background-color` بقيمة `rgb(127, 255, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(127, 255, 0)');
```

يجب أن يكون لقاعدة CSS لـ `.two` الخاصية `background-color` بقيمة `rgb(0, 127, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 127, 255)');
```

يجب أن يكون لقاعدة CSS لـ `.three` الخاصية `background-color` بقيمة `rgb(255, 0, 127)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(255, 0, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

.three {
  background-color: rgb(127, 0, 255);
}
--fcc-editable-region--

```
