---
id: 61b0a44a6b865738ba49b9fb
title: الخطوة 86
challengeType: 0
dashedName: step-86
---

# --description--

The last thing you'll do is add a slight shadow to each marker to make them look even more realistic.

الخاصية `box-shadow` تطبق ظلا أو أكثر حول العنصر. إليك بناء الجملة أولي:

```css
box-shadow: offsetX offsetY color;
```

إليك كيفية عمل قيم `offsetX` و `offsetY`:

* both `offsetX` and `offsetY` accept number values in `px` and other CSS units
* إن كانت قيمة `offsetX` موجبة سيتحرك الظل إلى اليمين أما قيمة سالبة تحركه إلى اليسار
* إن كانت قيمة `offsetY` موجبة سيتحرك الظل إلى الأسفل أما قيمة سالبة تحركه إلى فوق
* إذا كنت تريد قيمة صفر (`0`) لأي من `offsetX` أو `offsetY` أو كليهما، لا تحتاج إلى إضافة وحدة (unit). كل متصفح يفهم أن الصفر يعني لا تغيير.

يتم تحديد طول وعرض الظل حسب طول وعرض العنصر الذي طبق عليه. يمكنك أيضًا استخدام القيمة الاختيارية `spreadRadius` لتوسيع نطاق الظل. المزيد عن ذلك لاحقا.

ابدأ بإضافة ظل بسيط إلى marker الأحمر.

في قاعدة CSS المسمى `.red`، قم بإضافة خاصية `box-shadow` مع القيم `5px` ألى `offsetX` و `5px` ألى `offsetY` و `red` ألى `color`.

# --hints--

قاعدة CSS المسمى `.red`، يجب أن تحتوي علي الخاصية المختصرة `box-shadow` مع القيمة `5px 5px red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.red')?.boxShadow === 'red 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

--fcc-editable-region--
.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}
--fcc-editable-region--

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
