---
id: 587d8249367417b2b2512c42
title: Kifuatilia masuala
challengeType: 4
forumTopicId: 301569
dashedName: issue-tracker
---

# --description--

Jenga programu ya JavaScript ya full-stack ambayo inaendana kiutendaji na hii: <a href="https://issue-tracker.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://issue-tracker.freecodecamp.rocks/</a>. Kufanya kazi kwenye mradi huu kutahusisha kuandika msimbo wako kwa kutumia mojawapo ya njia zifuatazo:

-   Nakili <a href="https://github.com/freeCodeCamp/boilerplate-project-issuetracker/" target="_blank" rel="noopener noreferrer nofollow">hii repo ya GitHub</a> na ukamilishe mradi wako kwa ndani ya kompyuta yako.
-   Tumia mtengenezaji wa tovuti unayempenda kukamilisha mradi. Hakikisha unajumuisha mafaili yote kutoka kwenye repo yetu ya GitHub.

# --instructions--

-   Kamilisha njia za data zinazohitajika katika `/routes/api.js`
-   Unda majaribio yote ya kiutendaji katika `tests/2_functional-tests.js`
-   Nakili faili la `sample.env` kwenda `.env` na weka vigezo ipasavyo
-   Ili kuendesha majaribio moja kwa moja, ongeza `NODE_ENV=test` katika faili lako la `.env`
-   Ili kuendesha majaribio kwenye konsoli, tumia amri `npm run test`

Andika majaribio yafuatayo katika `tests/2_functional-tests.js`:

-   Unda suala na kila sehemu: ombi la POST kwenda `/api/issues/{project}`
-   Unda suala kwa kutumia sehemu zinazohitajika tu: ombi la POST kwenda `/api/issues/{project}`
-   Unda suala ukiwa na sehemu zinazohitajika zilizokosekana: ombi la POST kwenda `/api/issues/{project}`
-   Tazama masuala kwenye mradi: ombi la GET kwenda `/api/issues/{project}`
-   Tazama masuala kwenye mradi ukiwa na kichujio kimoja: ombi la GET kwenda `/api/issues/{project}`
-   Tazama masuala kwenye mradi ukiwa na vichujio vingi: ombi la GET kwenda `/api/issues/{project}`
-   Sasisha sehemu moja kwenye suala: ombi la PUT kwenda `/api/issues/{project}`
-   Sasisha sehemu nyingi kwenye suala: ombi la PUT kwenda `/api/issues/{project}`
-   Sasisha suala ukiwa na `_id` iliyokosekana: ombi la PUT kwenda `/api/issues/{project}`
-   Sasisha suala bila sehemu zozote za kusasisha: ombi la PUT kwenda `/api/issues/{project}`
-   Sasisha suala ukiwa na `_id` isiyo halali: ombi la PUT kwenda `/api/issues/{project}`
-   Futa suala: ombi la DELETE kwenda `/api/issues/{project}`
-   Futa suala ukiwa na `_id` isiyo halali: ombi la DELETE kwenda `/api/issues/{project}`
-   Futa suala ukiwa na `_id` iliyokosekana: ombi la DELETE kwenda `/api/issues/{project}`

# --hints--

Unaweza kutumia mradi wako mwenyewe, si lazima URL ya mfano.

```js
  assert(!/.*\/issue-tracker\.freecodecamp\.rocks/.test(code));
```

Unaweza kutuma ombi la `POST` kwenda `/api/issues/{projectname}` ukiwa na data ya fomu yenye sehemu zinazohitajika `issue_title`, `issue_text`, `created_by`, na hiari `assigned_to` na `status_text`.

```js
  try {
    let test_data = {
      issue_title: 'Faux Issue Title',
      issue_text: 'Functional Test - Required Fields Only',
      created_by: 'fCC'
    };
    const response = await fetch(code + '/api/issues/fcc-project', {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify(test_data)
    });
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.json();
    assert.isObject(data);
    assert.nestedInclude(data, test_data);
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

Ombi la `POST` kwenda `/api/issues/{projectname}` litarudisha kitu kilichoundwa, na lazima lijumuishe sehemu zote zilizowasilishwa. Sehemu za hiari zilizotengwa zitarudishwa kama mfuatano wa herufi tupu. Zaidi ya hayo, jumuisha `created_on` (tarehe/muda), `updated_on` (tarehe/muda), `open` (thamani ya kweli au siyo kweli, `true` kwa wazi - thamani ya msingi, `false` kwa kufungwa), na `_id`.

```js
  try {
    let test_data = {
      issue_title: 'Faux Issue Title 2',
      issue_text: 'Functional Test - Every field filled in',
      created_by: 'fCC',
      assigned_to: 'Chai and Mocha'
    };
    const response = await fetch(code + '/api/issues/fcc-project', {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify(test_data)
    });
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.json();
    assert.isObject(data);
    assert.nestedInclude(data, test_data);
    assert.property(data, 'created_on');
    assert.isNumber(Date.parse(data.created_on));
    assert.property(data, 'updated_on');
    assert.isNumber(Date.parse(data.updated_on));
    assert.property(data, 'open');
    assert.isBoolean(data.open);
    assert.isTrue(data.open);
    assert.property(data, '_id');
    assert.isNotEmpty(data._id);
    assert.property(data, 'status_text');
    assert.isEmpty(data.status_text);
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

Ikiwa utatuma ombi la `POST` kwenda `/api/issues/{projectname}` bila sehemu zinazohitajika, itarudisha kosa la `{ error: 'required field(s) missing' }`

```js
  try {
    let test_data = { created_by: 'fCC' };
    const response = await fetch(code + '/api/issues/fcc-project', {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ created_by: 'fCC' })
    });
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.json();
    assert.isObject(data);
    assert.property(data, 'error');
    assert.equal(data.error, 'required field(s) missing');
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

Unaweza kutuma ombi la `GET` kwenda `/api/issues/{projectname}` kwa safu ya data ya masuala yote ya `projectname` husika, ukiwa na sehemu zote zilizopo kwa kila suala.

```js
  try {
    let test_data = { issue_text: 'Get Issues Test', created_by: 'fCC' };
    const url =
      code +
      '/api/issues/get_issues_test_' +
      Date.now().toString().substring(7);
    const response1 = await fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify(Object.assign(test_data, { issue_title: 'Faux Issue 1' }))
    });
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.isObject(data1);
    const response2 = await fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify(Object.assign(test_data, { issue_title: 'Faux Issue 2' }))
    });
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.isObject(data2);
    const response3 = await fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify(Object.assign(test_data, { issue_title: 'Faux Issue 3' }))
    });
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.isObject(data3);
    const getResponse = await fetch(url);
    if (!getResponse.ok) {
      throw Error(await getResponse.text());
    }
    const getIssues = await getResponse.json();
    assert.isArray(getIssues);
    assert.lengthOf(getIssues, 3);
    let re = new RegExp('Faux Issue \\d');
    getIssues.forEach((issue) => {
      assert.property(issue, 'issue_title');
      assert.match(issue.issue_title, re);
      assert.property(issue, 'issue_text');
      assert.property(issue, 'created_by');
      assert.property(issue, 'assigned_to');
      assert.property(issue, 'status_text');
      assert.property(issue, 'open');
      assert.property(issue, 'created_on');
      assert.property(issue, 'updated_on');
      assert.property(issue, '_id');
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

Unaweza kutuma ombi la `GET` kwenda `/api/issues/{projectname}` na kuchuja ombi kwa pia kupitisha sehemu yoyote na thamani kama hoja ya utafutaji ya URL (yaani `/api/issues/{project}?open=false`). Unaweza kupitisha jozi moja au zaidi za sehemu/thamani kwa wakati mmoja.

```js
  try {
    let test_data = {
      issue_title: 'To be Filtered',
      issue_text: 'Filter Issues Test'
    };
    const url =
      code +
      '/api/issues/get_issues_test_' +
      Date.now().toString().substring(7);
    const response1 = await fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify(Object.assign(test_data, { created_by: 'Alice', assigned_to: 'Bob' }))
    });
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    const response2 = await fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify(Object.assign(test_data, { created_by: 'Alice', assigned_to: 'Bob' }))
    });
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    const response3 = await fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify(Object.assign(test_data, { created_by: 'Alice', assigned_to: 'Eric' }))
    });
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    const response4 = await fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify(Object.assign(test_data, { created_by: 'Carol', assigned_to: 'Eric' }))
    });
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    const getSingleResponse = await fetch(url + '?created_by=Alice');
    if (!getSingleResponse.ok) {
      throw Error(await getSingleResponse.text());
    }
    const getSingle = await getSingleResponse.json();
    assert.isArray(getSingle);
    assert.lengthOf(getSingle, 3);
    const getMultipleResponse = await fetch(url + '?created_by=Alice&assigned_to=Bob');
    if (!getMultipleResponse.ok) {
      throw Error(await getMultipleResponse.text());
    }
    const getMultiple = await getMultipleResponse.json();
    assert.isArray(getMultiple);
    assert.lengthOf(getMultiple, 2);
    const copyId = getMultiple[0]._id;
    const getByIdResponse = await fetch(url + `?_id=${copyId}`);
    if (!getByIdResponse.ok) {
      throw Error(await getByIdResponse.text());
    }
    const getById = await getByIdResponse.json();
    assert.isArray(getById);
    assert.lengthOf(getById, 1);
    assert.equal(getById[0]._id, copyId, 'should be able to query a document by _id')
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

Unaweza kutuma ombi la `PUT` kwenda `/api/issues/{projectname}` ukiwa na `_id` na sehemu moja au zaidi za kusasisha. Kwa mafanikio, sehemu ya `updated_on` inapaswa kusasishwa, na itarudisha `{  result: 'successfully updated', '_id': _id }`.

```js
  try {
    let initialData = {
      issue_title: 'Issue to be Updated',
      issue_text: 'Functional Test - Put target',
      created_by: 'fCC'
    };
    const url = code + '/api/issues/fcc-project';
    const createResponse = await fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify(initialData)
    });
    if (!createResponse.ok) {
      throw Error(await createResponse.text());
    }
    const itemToUpdate = await createResponse.json();
    const updateResponse = await fetch(url, {
      method: 'PUT',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ _id: itemToUpdate._id, issue_text: 'New Issue Text' })
    });
    if (!updateResponse.ok) {
      throw Error(await updateResponse.text());
    }
    const updateSuccess = await updateResponse.json();
    assert.isObject(updateSuccess);
    assert.deepEqual(updateSuccess, {
      result: 'successfully updated',
      _id: itemToUpdate._id
    });
    const getUpdatedResponse = await fetch(url + '?_id=' + itemToUpdate._id);
    if (!getUpdatedResponse.ok) {
      throw Error(await getUpdatedResponse.text());
    }
    const getUpdatedId = await getUpdatedResponse.json();
    assert.isArray(getUpdatedId);
    assert.isObject(getUpdatedId[0]);
    assert.isAbove(
      Date.parse(getUpdatedId[0].updated_on),
      Date.parse(getUpdatedId[0].created_on)
    );
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

Wakati ombi la `PUT` lililotumwa kwenda `/api/issues/{projectname}` halijajumuisha `_id`, thamani ya return ni `{ error: 'missing _id' }`.

```js
  try {
    const url = code + '/api/issues/fcc-project';
    const response = await fetch(url, { method: 'PUT' });
    if (!response.ok) {
      throw Error(await response.text());
    }
    const badUpdate = await response.json();
    assert.isObject(badUpdate);
    assert.property(badUpdate, 'error');
    assert.equal(badUpdate.error, 'missing _id');
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

Wakati ombi la `PUT` lililotumwa kwenda `/api/issues/{projectname}` halijajumuisha sehemu za kusasisha, thamani ya return ni `{ error: 'no update field(s) sent', '_id': _id }`. Kwa kosa lolote lingine, thamani ya return ni `{ error: 'could not update', '_id': _id }`.

```js
  try {
    const url = code + '/api/issues/fcc-project';
    const response1 = await fetch(url, {
      method: 'PUT',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ _id: '5f665eb46e296f6b9b6a504d' })
    });
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const badUpdate = await response1.json();
    assert.deepEqual(badUpdate, {
      error: 'no update field(s) sent',
      _id: '5f665eb46e296f6b9b6a504d'
    });
    const response2 = await fetch(url, {
      method: 'PUT',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ _id: '5f665eb46e296f6b9b6a504d', issue_text: 'New Issue Text' })
    });
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const badIdUpdate = await response2.json();
    assert.deepEqual(badIdUpdate, {
      error: 'could not update',
      _id: '5f665eb46e296f6b9b6a504d'
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

Unaweza kutuma ombi la `DELETE` kwenda `/api/issues/{projectname}` ukiwa na `_id` kufuta suala. Ikiwa hakuna `_id` lililotumwa, thamani ya return ni `{ error: 'missing _id' }`. Kwa mafanikio, thamani ya return ni `{ result: 'successfully deleted', '_id': _id }`. Kwa kushindwa, thamani ya return ni `{ error: 'could not delete', '_id': _id }`.

```js
  try {
    let initialData = {
      issue_title: 'Issue to be Deleted',
      issue_text: 'Functional Test - Delete target',
      created_by: 'fCC'
    };
    const url = code + '/api/issues/fcc-project';
    const createResponse = await fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify(initialData)
    });
    if (!createResponse.ok) {
      throw Error(await createResponse.text());
    }
    const itemToDelete = await createResponse.json();
    assert.isObject(itemToDelete);
    const deleteResponse = await fetch(url, {
      method: 'DELETE',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ _id: itemToDelete._id })
    });
    if (!deleteResponse.ok) {
      throw Error(await deleteResponse.text());
    }
    const deleteSuccess = await deleteResponse.json();
    assert.isObject(deleteSuccess);
    assert.deepEqual(deleteSuccess, {
      result: 'successfully deleted',
      _id: itemToDelete._id
    });
    const noIdResponse = await fetch(url, { method: 'DELETE' });
    if (!noIdResponse.ok) {
      throw Error(await noIdResponse.text());
    }
    const noId = await noIdResponse.json();
    assert.isObject(noId);
    assert.deepEqual(noId, { error: 'missing _id' });
    const badIdDeleteResponse = await fetch(url, {
      method: 'DELETE',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ _id: '5f665eb46e296f6b9b6a504d', issue_text: 'New Issue Text' })
    });
    if (!badIdDeleteResponse.ok) {
      throw Error(await badIdDeleteResponse.text());
    }
    const badIdDelete = await badIdDeleteResponse.json();
    assert.isObject(badIdDelete);
    assert.deepEqual(badIdDelete, {
      error: 'could not delete',
      _id: '5f665eb46e296f6b9b6a504d'
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

Majaribio yote 14 ya kiutendaji yamekamilika na kupita.

```js
  try {
    const response = await fetch(code + '/_api/get-tests');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const getTests = await response.json();
    assert.isArray(getTests);
    assert.isAtLeast(getTests.length, 14, 'At least 14 tests passed');
    getTests.forEach((test) => {
      assert.equal(test.state, 'passed', 'Test in Passed State');
      assert.isAtLeast(
        test.assertions.length,
        1,
        'At least one assertion per test'
      );
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```
