---
id: 669aff9f5488f1bea056416d
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本工作坊中，你将通过创建一个专门介绍猫 Whiskers 先生的博客页面来练习使用语义化 HTML。

开始项目，添加 `<!DOCTYPE html>`，添加一个 `html` 元素，其 `lang` 属性值为 `en`。

请记住，你在上一模块中学过如何创建类似的基本 HTML 模板。

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

# --hints--

你应该有 `<!DOCTYPE html>`。

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

你应该有一个 `html` 开始标签，语言设置为英语。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

你应该有一个 `html` 结束标签。

```js
assert.match(code, /<\/html>/i);
```

你的 `DOCTYPE` 应放在 `html` 元素之前。

```js
assert.match(code, /<!DOCTYPE\s+html>[.\n\s]*<html\s+lang\s*=\s*('|")en\1\s*>/im)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
