---
id: 669fcb06c3034a39f5431a38
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

为了准备创建一些实际的内容，在 `head` 元素下面添加一个 `body` 元素。

# --hints--

你应该有一个 `<body>` 开始标签。

```js
assert.match(code, /<body>/i);
```

你应该有一个 `</body>` 结束标签。

```js
assert.match(code, /<\/body>/i);
```

你不应该改变你的 `head` 元素。 确认你没有删除结束标签。

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

你的 `body` 元素应该在你的 `head` 元素之后。

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>

  --fcc-editable-region--
</html>
```
