---
id: 669fd003cb89ee3c2402e041
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

在这个介绍性内容中，需要展示一张 Whiskers 先生的图片，并配上说明。

在 `h1` 元素下方，首先添加一个 `figure` 元素。

在 `figure` 元素中，添加一个 `img` 元素。

`img` 的 `src` 属性值应为 `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`，`alt` 文本应为 `"a cat in the garden"`。

在你的 `img` 元素下方，添加一个 `figcaption</code`，文本为 `Mr. Whiskers in the Garden`。

# --hints--

你的 `figure` 元素应该有一个开始标签。

```js
assert.match(code, /<figure\>/);
```

你的 `figure` 元素应该有一个结束标签。

```js
assert.match(code, /<\/figure\>/);
```

你的 `figure` 元素应该在 `header` 元素中。

```js
const figureElement = document.querySelector('figure');
assert.equal(figureElement?.parentElement.tagName, 'HEADER');
```

你应该有一个 `img` 元素。

```js
assert.exists(document.querySelector("img"));
```

你的 `img` 元素应在 `figure` 元素中。

```js
const imgElement = document.querySelector('img');
assert.strictEqual(imgElement?.parentElement.tagName, 'FIGURE');
```

你的 `img` 元素应该有一个 `src` 属性，设置为 `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`。

```js
const imgElement = document.querySelector('img');
const src = imgElement?.getAttribute('src');

assert.strictEqual(src, 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

你的 `img` 元素应该有一个 `alt` 属性，设置为 `"a cat in the garden"`。

```js
const imgElement = document.querySelector('img');
const alt = imgElement?.getAttribute('alt');

assert.strictEqual(alt, 'a cat in the garden');
```

你的 `figcaption` 元素应该有一个开始标签。

```js
assert.match(code, /<figcaption\>/);
```

你的 `figcaption` 元素应该有一个结束标签。

```js
assert.match(code, /<\/figcaption\>/);
```

你的 `figcaption` 元素应在 `figure` 元素中。

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.parentElement.tagName, 'FIGURE');
```

你的 `figcaption` 元素的文本应为 `Mr. Whiskers in the Garden`。

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.innerText.trim(), "Mr. Whiskers in the Garden")
```

你的 `figcaption` 应在 `img` 元素下方。 你把它们的顺序写错了。

```js
const figureElement = document.querySelector('figure');
assert.strictEqual(figureElement?.lastElementChild?.tagName, "FIGCAPTION")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
--fcc-editable-region--

--fcc-editable-region--
    </header>
  </body>
</html>
```
