---
id: 669fdc11c9b0263fe0814a7a
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

在每个 `li` 元素中，都需要有一个锚元素。

第一个锚元素的文本应为 `About`，`href` 属性值应为 `"#about"`。 `about` 前面的哈希符号代表一个 `id` 名称，稍后将添加到项目中。

第二个锚元素的文本应为 `Posts`，`href` 属性值应为 `"#posts"`。

第三个锚元素的文本应为 `Contact`，`href` 属性值应为 `"#contact"`。

# --hints--

你的第一个锚元素应在第一个 `li` 元素中包含 `About` 文本。

```js
const anchorElement = document.querySelector("ul li:first-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "About");
```

你的第一个锚元素应具有 `href` 属性，设置为 `"#about"`。

```js
const anchorElement = document.querySelector("ul li:first-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#about");
```

你的第二个锚元素应在第二个 `li` 元素中包含 `Posts` 文本。

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
assert.strictEqual(anchorElement?.innerText.trim(), "Posts");
```

你的第二个锚元素应具有 `href` 属性，设置为 `"#posts"`。

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#posts");
```

你的第三个锚元素应在第三个 `li` 元素中包含 `Contact` 文本。

```js
const anchorElement = document.querySelector("ul li:last-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "Contact");
```

你的第三个锚元素应具有 `href` 属性，设置为 `"#contact"`。

```js
const anchorElement = document.querySelector("ul li:last-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#contact");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
--fcc-editable-region--
          <li></li>
          <li></li>
          <li></li>
--fcc-editable-region--
        </ul>
      </nav>
    </header>
  </body>
</html>
```
