---
id: 669fcb5106eeb13ab7afcd09
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

In this workshop, you will practice working with semantic HTML by building a blog page dedicated to Mr. Whiskers the cat.

你要创建的第一个部分是<dfn>页头</dfn>。

`header` 元素用于表示介绍性内容，如页面导航和其他介绍性信息。

下面是一个使用 `header` 元素的示例：

```html
<header>
  <h1>Main Page Title Goes Here</h1>
  <img src="example-logo.png" alt="Example logo" />
</header>
```

在 `body` 元素内，添加一个 `header` 元素。

# --hints--

你应该有一个 `header` 开始标签。

```js
assert.match(code, /<header>/i);
```

你应该有一个 `header` 结束标签。

```js
assert.match(code, /<\/header>/i);
```

你的 `header` 元素应该在 `body` 元素中。

```js
assert.match(code, /<\/header>[.\n\s]*<\/body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```
