---
id: 669fc7e141e4703748c558bf
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `html` 元素中，添加一个 `head` 元素。

# --hints--

应该有一个 `head` 开始标签。

```js
assert.match(code, /<head>/i);
```

应该有一个 `head` 结束标签。

```js
assert.match(code, /<\/head>/i);
```

你的 `head` 开始标签应在 `head` 结束标签之前。

```js
assert.match(code, /<head>[.\n\s]*<\/head>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
