---
id: 66a334a1a7cca6354999f9bf
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

对于第一篇博文，将使用 <dfn>article</dfn> 元素。

`article` 元素表示网页中的自包含内容。

```html
<article>
  <h1>Example heading</h1>
  <p>Example article text</p>
</article>
```

在 `h2` 元素下方，添加一个 `article` 元素。

在 `article` 元素中，添加一个 `h3` 元素，其文本为 `Mr. Whiskers' First Day Home`。

这里使用 `h3` 的原因是，为标题元素保持适当的结构层级非常重要。 由于博文副标题是一个 `h2` 元素，层级结构中的下一级是 `h3` 。

# --hints--

你的代码应该有一个 `article` 开始标签。

```js
assert.match(code, /<article>/i);
```

你的代码应该有一个 `article` 结束标签。

```js
assert.match(code, /<\/article>/i);
```

你的 `article` 元素应在 `section` 元素中，`id` 设置为 `posts`。

```js
const articleElement = document.querySelector('#posts article');
assert.strictEqual(articleElement?.parentElement.tagName, 'SECTION');
```

你的代码应该有一个 `h3` 开始标签。

```js
assert.match(code, /<h3>/i);
```

你的代码应该有一个 `h3` 结束标签。

```js
assert.match(code, /<\/h3>/i);
```

你的 `h3` 元素的文本应为 `Mr. Whiskers' First Day Home`。

```js
const h3Element = document.querySelector('h3');
assert.strictEqual(h3Element?.innerText.trim(), "Mr. Whiskers' First Day Home");
```

你的 `h3` 元素应在 `article` 元素中。

```js
const h3Element = document.querySelector('#posts article h3');
assert.strictEqual(h3Element?.parentElement.tagName, 'ARTICLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
