---
id: 66a49f685961e997e337cab1
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

`header` 将负责显示博客的主标题、图像和页面导航。

在 `header` 元素中，添加一个 `h1`，其文本为 `Welcome to Mr. Whiskers' Blog Page!`。

# --hints--

你应该有一个 `h1` 开始标签。

```js
assert.match(code, /<h1>/i);
```

你应该有一个 `h1` 结束标签。

```js
assert.match(code, /<\/h1>/i);
```

你的 `h1` 元素的文本应为 `Welcome to Mr. Whiskers' Blog Page!`。 仔细检查是否有拼写错误。

```js
const h1Text = document.querySelector("h1")?.innerText.trim();
assert.strictEqual(h1Text, "Welcome to Mr. Whiskers' Blog Page!");
```

你的 `h1` 元素应该嵌套在你的 `header` 元素中。

```js
const h1Element = document.querySelector("h1");
assert.strictEqual(h1Element?.parentElement.tagName, 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
  </body>
</html>
```
