---
id: 68de40f53fcaf106f51cc20f
title: 构建一个超文本标记语言音频和视频播放器
challengeType: 25
dashedName: build-an-html-audio-and-video-player
demoType: onClick
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个 `h1` 元素，作为页面的主标题。
2. 你应该有两个 `section` 元素。
3. 在第一个 `section` 元素内，你应该有一个 `h2` 元素作为正在播放的视频标题。
4. 在 `h2` 元素下方，你应该有一个带有 `controls` 和 `width` 属性的 `video` 元素。`width` 属性应设置为 `640`。
5. 在 `video` 元素内部，应该有一个 `source` 元素，带有指向视频文件的 `src` 属性和 `type` 属性。
   - 你可以使用 `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`。
6. 在第二个 `section` 元素内，你应该有一个 `h2` 元素作为正在播放歌曲的标题。
7. 在 `h2` 元素下方，你应该有一个带有 `controls` 和 `loop` 属性的 `audio` 元素，以及一个指向音频文件的 `src` 属性。
   - 你可以使用 `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`。
   - 或者 `https://cdn.freecodecamp.org/curriculum/js-music-player/we-are-going-to-make-it.mp3`。

# --hints--

你应该有一个 `h1` 元素，作为页面的主标题。

```js
const h1 = document.querySelector('h1');
assert.exists(h1);
```

你的 `h1` 不应为空。

```js
const h1 = document.querySelector('h1');
assert.isNotEmpty(h1.textContent.trim());
```

你应该只有一个 `h1` 元素

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你应该有两个 `section` 元素。

```js
assert.lengthOf(document.querySelectorAll('section'), 2);

```

在第一个 `section` 元素内，你应该有一个 `h2` 元素作为正在播放的视频标题。

```js
const section = document.querySelector('section');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

在 `h2` 元素下方，你应该有一个 `video` 元素。

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
assert.exists(video);
```

`video` 元素应具有 `controls` 属性。

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const controlsAttr = video.controls;
assert.isTrue(controlsAttr);
```

`video` 元素应具有设置为 `640` 的 `width` 属性。

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const widthAttr = video.getAttribute('width');
assert.equal(widthAttr, '640');
```

在 `video` 元素内，你应该有一个 `source` 元素。

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
assert.exists(source);
```

该 `source` 元素应该有一个指向视频文件的 `src` 属性。

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const srcAttr = source.getAttribute('src');
assert.match(srcAttr, /\.(mp4|webm|ogg|avi|mov)$/i);
```

`source` 元素应具有 `type` 属性

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
assert.isNotEmpty(typeAttr);
```

`type` 属性应匹配 `src` 属性的文件扩展名的媒体类型。

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
const srcAttr = source.getAttribute('src');
const srcFileExt = srcAttr.match(/[^\.]+$/g)[0];
const fileExtToMediaTypes = {
  mp4: ['video/mp4'],
  webm: ['video/webm'],
  ogg: ['video/ogg'],
  avi: ['video/avi', 'video/vnd.avi', 'video/msvideo', 'video/x-msvideo'],
  mov: ['video/quicktime']
};
const srcMediaTypes = fileExtToMediaTypes[srcFileExt];
assert.oneOf(typeAttr, srcMediaTypes);
```

在第二个 `section` 元素内，你应该有一个 `h2` 元素作为正在播放歌曲的标题。

```js
const section = document.querySelector('section:nth-of-type(2)');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

你所有的 `h2` 元素都应包含一些文本。

```js
const h2s = document.querySelectorAll('h2');
assert.isNotEmpty(h2s);
h2s.forEach((h2) => assert.isNotEmpty(h2.textContent.trim()));
```

在 `h2` 元素下方，你应该有一个 `audio` 元素。

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
assert.exists(audio);
```

`audio` 元素应具有 `controls` 属性。

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const controlsAttr = audio.controls;
assert.isTrue(controlsAttr);
```

`audio` 元素应具有 `loop` 属性。

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const loopAttr = audio.loop;
assert.isTrue(loopAttr);
```

`audio` 元素应具有指向音频 `file` 的 `src` 属性。

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const srcAttr = audio.getAttribute('src');
assert.match(srcAttr, /\.(mp3|wav|ogg|aac|flac)$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>
  <main>
    <h1>HTML Audio and Video Lab</h1>

    <section>
      <h2>What is the map method and how does it work?</h2>
      <video controls width="640">
        <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4" />
      </video>
    </section>

    <section>
      <h2>Sailing Away</h2>
      <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" controls loop>
      </audio>
    </section>

  </main>
</body>
</html>
```
