---
id: 669e2f60e83c011754f711f9
title: 여행사 페이지 만들기
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `DOCTYPE` 선언이 있어야 합니다.
1. `lang`이 `en`으로 설정된 `html` 요소가 있어야 합니다.
1. `head` 요소 안에 `charset`이 `utf-8`로 설정된 `meta` 빈 요소와, 텍스트가 `Travel Agency Page`인 `title`이 있어야 합니다.
1. `head` 요소 안에 SEO를 위한 웹사이트의 간단한 설명이 들어 있는 `meta` 태그가 있어야 합니다.
1. 여행 목적지를 소개하는 `h1` 요소가 있어야 합니다. 이 `h1` 요소는 페이지에 하나만 있어야 합니다.
1. `h1` 요소 아래에 여행 기회를 소개하는 단락이 있어야 합니다.
1. 텍스트가 `Packages`인 `h2` 요소가 있어야 합니다.
1. 다양한 패키지를 간단히 소개하는 `p` 요소가 있어야 합니다.
1. 두 개의 항목이 있는 순서 없는 목록 요소가 있어야 합니다. 두 목록 항목의 텍스트는 각각 `Group Travels`와 `Private Tours`여야 하며, 각 항목의 텍스트는 앵커 요소로 감싸야 합니다.
1. 텍스트가 `Top Itineraries`인 `h2` 요소가 있어야 합니다.
1. 최소 세 개의 `figure` 요소가 있어야 하며, 각 요소에는 앵커 요소와 그림의 이미지를 설명하는 `figcaption` 요소가 포함되어야 합니다.
1. 세 개의 앵커 요소에는 적절한 `alt` 속성과 유효한 이미지로 설정된 `src` 속성을 가진 `img` 요소가 콘텐츠로 포함되어야 합니다. 원한다면 `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`, `https://cdn.freecodecamp.org/curriculum/labs/alps.jpg`, `https://cdn.freecodecamp.org/curriculum/labs/sea.jpg`를 사용할 수 있습니다.
1. 다섯 개의 모든 앵커 요소에는 `href` 속성 값이 `https://www.freecodecamp.org/learn`이고, `target` 속성 값이 `_blank`여야 합니다.

# --hints--

여행사 페이지에는 `!DOCTYPE html` 선언이 있어야 합니다.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

`html`가 `lang`로 설정된 `en` 요소가 있어야 합니다.

```js
assert.match(code, /<\s*html(?=[\s>])[^>]*\slang\s*=\s*('|")en\1[^>]*>/i);
```

`html` 요소는 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/html\s*>/i);
```

`head` 요소 안에 `html` 요소가 있어야 합니다.

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>/i
);
```

`head` 요소에 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/head\s*>/i);
```

`head` 요소 안에 `title` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelector('head title'));
```

`title` 요소에 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/title\s*>/i);
```

`head` 요소 안에 두 개의 `meta` 요소가 있어야 합니다.

```js
assert.isAtLeast(document.querySelectorAll('head meta').length, 2);
```

한 `meta` 요소는 `description` 값을 가진 `name` 속성과 비어 있지 않은 `content` 속성을 가져야 합니다.

```js
assert(
  Array.from(document.querySelectorAll('head meta')).some(
    (meta) =>
      meta.getAttribute('name')?.toLowerCase() === 'description' &&
      meta.getAttribute('content')?.trim()
  )
);
```

`meta` 요소 중 하나는 `charset` 속성이 `utf-8`로 설정되어야 합니다.

```js
const charsetMeta = document.querySelector('head meta[charset]');
assert.equal(charsetMeta?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

당신의 `title` 요소에는 여행사 이름이 있어야 합니다.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

`body` 요소 안에 `html` 요소가 있어야 합니다.

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>(?:(?!<\/html\s*>)[\s\S])*<\s*body\s*>/i
);
```

`body` 요소에 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/body\s*>/i);
```

여행지를 소개하는 `h1` 요소가 있어야 합니다.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

`h1` 요소는 하나만 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

`h1` 요소 아래에 `p` 요소가 있어야 합니다.

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

첫 번째 `p` 요소는 여행 기회를 소개해야 합니다.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

첫 번째 `h2` 요소에는 `Packages`라는 텍스트가 있어야 합니다.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText.trim(), 'Packages');
```

첫 번째 `h2` 요소 아래에 `p` 요소가 있어야 합니다.

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

두 번째 `p` 요소는 다양한 패키지를 간략히 소개해야 합니다.

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

두 번째 `p` 요소 아래에 순서 없는 목록 요소가 있어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

순서 없는 목록에는 두 개의 항목이 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

두 목록 항목 모두에 앵커 요소가 포함되어야 합니다.

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

첫 번째 목록 항목의 앵커 요소는 `Group Travels` 텍스트를 감싸야 합니다.

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText.trim(), 'Group Travels');
```

두 번째 목록 항목의 앵커 요소는 `Private Tours` 텍스트를 감싸야 합니다.

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText.trim(), 'Private Tours');
```

순서 없는 목록 다음에 `h2` 요소가 있어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

두 번째 `h2` 요소에는 `Top Itineraries`라는 텍스트가 있어야 합니다.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText.trim(), 'Top Itineraries');
```

최소 세 개의 `figure` 요소가 있어야 합니다.

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

각 `figure` 요소는 첫 번째 중첩 요소로 앵커 요소를 포함해야 합니다.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

각 `figure` 요소는 두 번째 중첩 요소로 `figcaption` 요소를 포함해야 합니다.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

각 `figcaption`에는 텍스트가 포함되어야 합니다.

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

각 `figure` 요소에 중첩된 `a` 요소는 이미지를 포함해야 합니다.

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

각 `img` 요소에는 유효한 `src` 속성이 있어야 합니다.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

각 `img` 요소에는 적절한 값을 가진 `alt` 속성이 있어야 합니다.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

각 `a` 요소에는 `https://www.freecodecamp.org/learn` 값을 가진 `href` 속성이 있어야 합니다. 목록 항목의 링크도 잊지 마세요.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

각 `a` 요소에는 `_blank` 값을 가진 `target` 속성이 있어야 합니다. 목록 항목의 링크도 잊지 마세요.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```
