---
id: 68e707257ea91c2e12d489fb
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, practicar치s c칩mo agregar estilos personalizados a los botones de opci칩n construyendo un formulario para la conferencia entre padres y maestros. Se te ha proporcionado la plantilla b치sica de HTML.

Comienza agregando un elemento `main` con una clase llamada `container`.

# --hints--

Debes tener un elemento `main`.

```js
assert.exists(document.querySelector("main"));
```

Tu elemento `main` debe tener una clase llamada `container`.

```js
assert.exists(document.querySelector("main.container"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
