---
id: 68e7e729f6d14f69a23786ef
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

A continuación, dentro de tu elemento `main`, agrega un elemento `h1` con el texto `Parent Teacher Conference Form` y las clases `title` y `center`. 

# --hints--

Debería tener un elemento `h1`.

```js
assert.exists(document.querySelector("h1"));
```

Tu elemento `h1` debe tener el texto `Parent Teacher Conference Form`.

```js
assert.strictEqual(document.querySelector("h1")?.innerText, "Parent Teacher Conference Form");
```

Tu elemento `h1` debe tener las clases `title` y `center`.

```js
assert.exists(document.querySelector("h1.title.center"));
```

Tu elemento `h1` debe estar dentro del elemento `main`.

```js
assert.exists(document.querySelector("main h1.title.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  --fcc-editable-region--
  <main class="container">
    
  </main>
  --fcc-editable-region--
</body>

</html>
```
