---
id: 68ea8666b8b473455f5b33ef
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

La primera sección del formulario se centrará en la información del estudiante, como su nombre y grado.

Dentro del elemento `form`, agrega un elemento `fieldset`. Dentro de ese elemento `fieldset`, agrega un elemento `legend` con el texto `Student Information`. 

Debajo del elemento `legend`, agrega un elemento `label` con el texto `Full Name: ` y un atributo `for` con el valor de `student-name`.


# --hints--

Debería tener un elemento `fieldset`.

```js
assert.exists(document.querySelector("fieldset"));
```

Tu elemento `fieldset` debería estar dentro del elemento `form`.

```js
assert.exists(document.querySelector("form fieldset"));
```

Debes tener un elemento `legend`.

```js
assert.exists(document.querySelector("legend"));
```

Tu elemento `legend` debería estar dentro del elemento `fieldset`.

```js
assert.exists(document.querySelector("fieldset legend"));
```

Tu elemento `legend` debe tener el texto `Student Information`.

```js
assert.strictEqual(document.querySelector("fieldset legend")?.innerText, "Student Information");
```

Deberías tener un elemento `label`.

```js
assert.exists(document.querySelector("label"));
```

Tu elemento `label` debe tener el texto `Full Name: `. No olvides el espacio después de los dos puntos.

```js
assert.match(code, /Full\s+Name:\s+<\/label>/);
```

Tu elemento `label` debe tener un atributo `for`.

```js
assert.isTrue(document.querySelector("label")?.hasAttribute("for"));
```

El atributo `for` del elemento `label` debe tener el valor `student-name`.

```js
assert.strictEqual(document.querySelector("label")?.getAttribute("for"), "student-name");
```

Tu elemento `label` debe estar debajo del elemento `legend`.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const legendEl = document.querySelector("fieldset legend");
const labelEl = document.querySelector("label");

assert.isBelow(fieldsetChildren.indexOf(legendEl), fieldsetChildren.indexOf(labelEl));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    --fcc-editable-region--
    <form>

      
    </form>
    --fcc-editable-region--
  </main>
</body>

</html>
```
