---
id: 68ea975c7cbbb894a386e402
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

El siguiente paso es agregar los elementos `label` e `input` para el método de contacto por correo electrónico.

Comienza agregando una etiqueta con una clase de `"contact-method"` y un atributo `for` establecido en `"email"`. El texto de la etiqueta debe ser `Email: `.

Debajo de tu elemento `label`, agrega un botón de opción con los atributos `id` y `value` establecidos en `"email"`. El atributo `name` debe estar configurado en `"contact-method"` y la clase debe estar configurada en `"contact-method-radio-btn"`. 

Por último, asegúrate de que este botón de opción esté seleccionado por defecto.

# --hints--


Tu último fieldset debe contener un elemento `label`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

Este nuevo elemento `label` (para email) debe tener una clase `"contact-method"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

El nuevo atributo `for` del elemento `label` debe establecerse en `"email"`.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "email");
```

El elemento `label` debe tener el texto `Email: `. Asegúrate del espacio después de los dos puntos.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.match(label?.innerText, /Email:\s?/);
```

Debes tener un elemento `input` dentro del último `fieldset`.

```js
assert.exists(document.querySelector("fieldset:last-of-type input"));
```

Tu elemento `input` debe tener un atributo `type` establecido en `radio`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

El elemento `input` debe tener un atributo `id` establecido en `"email"`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("id"), "email");
```

El elemento `input` debe tener un atributo `value` establecido en `"email"`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("value"), "email");
```

El elemento `input` debe tener un atributo `name` establecido en `"contact-method"`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

El elemento `input` debe tener una clase de `"contact-method-radio-btn"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type input.contact-method-radio-btn"));
```

Tu botón `radio` debería estar `checked` por defecto.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.isTrue(input?.hasAttribute("checked"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Information: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      --fcc-editable-region--
      <fieldset>
        <legend>Preferred Contact Method</legend>
        
      </fieldset>
      --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
