---
id: 68eaa668d4d9fbe529a0e16b
title: Paso 14
challengeType: 0
dashedName: step-14
---

# --description--

Ahora, agrega un elemento `label` con el texto `Any specific concerns or topics you'd like to discuss?` y un atributo `for` establecido en `notes`.

Debajo de este `label`, agrega un elemento `textarea`. Establece su `id` en `notes`, `name` en `notes`, `rows` en `4` y `cols` en `50`.

# --hints--

Tu último `fieldset` debe contener un elemento `label`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

El elemento `label` debe tener el texto `Any specific concerns or topics you'd like to discuss?`.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.innerText, "Any specific concerns or topics you'd like to discuss?");
```

El atributo `for` del elemento `label` debe establecerse en `"notes"`.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "notes");
```

Tu último `fieldset` debe contener un elemento `textarea`.

```js
assert.exists(document.querySelector("fieldset:last-of-type textarea"));
```

Tu elemento `textarea` debe estar debajo del elemento `label`.

```js
const label = document.querySelector("fieldset:last-of-type label");
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(label?.nextElementSibling, textarea);
```

El elemento `textarea` debe tener un atributo `id` establecido en `"notes"`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("id"), "notes");
```

El elemento `textarea` debe tener un atributo `name` establecido en `"notes"`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("name"), "notes");
```

El elemento `textarea` debe tener un atributo `rows` configurado en `4`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("rows"), "4");
```

El elemento `textarea` debe tener un atributo `cols` configurado en `50`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("cols"), "50");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Information: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>

        --fcc-editable-region--
        
        --fcc-editable-region--

      </fieldset>
    </form>
  </main>
</body>

</html>
```
