---
id: 68ea8c330ed07b60879f3d38
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Ahora es momento de agregar los elementos del formulario para recopilar la información de la calificación del estudiante.

Comienza agregando otro elemento `label` con el texto `Student Grade: ` y el atributo `for` establecido en `"grade"`.

Luego, debajo de tu elemento `label`, agrega un elemento `input` con el atributo `type` establecido en `"number"`. Los atributos `name` e `id` deben estar establecidos en `"grade"`. El atributo `placeholder` debe estar establecido en `"E.g., 4"`. Por último, tu `input` de número debe ser obligatorio. 

# --hints--

Deberías tener un total de dos elementos `label` en la página.

```js
assert.lengthOf(document.querySelectorAll("label"), 2);
```

Tu segundo elemento `label` debe tener el texto `Student Grade: `. No olvides el espacio después de los dos puntos.

```js
assert.match(code, /Student\s+Grade:\s+<\/label>/);
```

Tu segundo elemento `label` debe tener un atributo `for`.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

El atributo `for` del segundo elemento `label` debe tener el valor `"grade"`.

```js
assert.strictEqual(document.querySelector("label:last-of-type")?.getAttribute("for"), "grade");
```

Debes tener dos elementos `input` en la página.

```js
assert.lengthOf(document.querySelectorAll("input"), 2);
```

Tu segundo elemento `input` debe estar debajo de tu segundo elemento `label`.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const inputEl = document.querySelector("fieldset input:last-of-type");
const labelEl = document.querySelector("label:last-of-type");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

Tu segundo elemento `input` debe tener un atributo `type` establecido en `"number"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("type"), "number");
```

Tu segundo elemento `input` debe tener un atributo `name` establecido en `"grade"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("name"), "grade");
```

Tu segundo elemento `input` debe tener un atributo `id` establecido en `"grade"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("id"), "grade");
```

Tu segundo elemento `input` debe tener un atributo `placeholder` establecido en `"E.g., 4"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., 4");
```

Tu segundo elemento `input` debe tener un atributo `required`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
