---
id: 68ea8fdf2c9b167307e7f671
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Ahora es momento de agregar los elementos del formulario para recopilar la información del padre.

Comienza agregando un elemento `label` con el texto `Parent/Guardian Name: ` y el atributo `for` establecido en `"parent-name"`.

Luego, debajo de tu elemento `label`, agrega un elemento `input` con el atributo `type` establecido en `"text"`. Los atributos `name` e `id` deben estar establecidos en `"parent-name"`. El atributo `placeholder` debe estar establecido en `"E.g., Nancy Doe"`. Por último, tu `input` debe ser obligatorio. 

# --hints--

Deberías tener un total de tres elementos `label` en la página.

```js
assert.lengthOf(document.querySelectorAll("label"), 3);
```

Tu tercer elemento `label` debe tener el texto `Parent/Guardian Name: `. No olvides el espacio después de los dos puntos.

```js
assert.match(code, /Parent\/Guardian\s+Name:\s+<\/label>/);
```

Tu tercer elemento `label` debe tener un atributo `for`.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

El atributo `for` del tercer elemento `label` debe tener el valor `"parent-name"`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type label")?.getAttribute("for"), "parent-name");
```

Debes tener tres elementos `input` en la página.

```js
assert.lengthOf(document.querySelectorAll("input"), 3);
```

Tu tercer elemento `input` debe estar debajo de tu tercer elemento `label`.

```js
const fieldsetChildren = [...document.querySelector("fieldset:last-of-type")?.children];
const inputEl = document.querySelector("fieldset:last-of-type input");
const labelEl = document.querySelector("fieldset:last-of-type label");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

Tu tercer elemento `input` debe tener un atributo `type` establecido en `"text"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

Tu tercer elemento `input` debe tener un atributo `name` establecido en `"parent-name"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("name"), "parent-name");
```

Tu tercer elemento `input` debe tener un atributo `id` establecido en `"parent-name"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("id"), "parent-name");
```

Tu tercer elemento `input` debe tener un atributo `placeholder` establecido en `"E.g., Nancy Doe"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Nancy Doe");
```

Tu tercer elemento `input` debe tener un atributo `required`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
