---
id: 68ea8a5db73b155708992687
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

El siguiente paso es agregar el elemento `input` asociado para la información del estudiante.

Comienza agregando un elemento `input` con un atributo `type` establecido en `"text"`. Luego agrega un atributo `name` y `id`, ambos establecidos en `student-name`. 

Luego, agrega un atributo `placeholder` con el valor `E.g., Jane Doe`. Y finalmente, agrega un atributo `required`. 

# --hints--

Debes tener un elemento `input`.

```js
assert.exists(document.querySelector("fieldset input"));
```

Tu elemento `input` debe tener un atributo `type` establecido en `"text"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

Tu elemento `input` debe tener un atributo `name` establecido en `"student-name"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("name"), "student-name");
```

Tu elemento `input` debe tener un atributo `id` establecido en `"student-name"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("id"), "student-name");
```

Tu elemento `input` debe tener un atributo `placeholder` configurado como `E.g., Jane Doe`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Jane Doe");
```

Tu elemento `input` debe tener el atributo `required`.

```js
const inputEl = document.querySelector("fieldset input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>

      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>


    </form>
  </main>
</body>

</html>
```
