---
id: 68eb2161d02d00c013fc44f9
title: Pasos 24
challengeType: 0
dashedName: step-24
---

# --description--

Para hacer que el texto de la `legend` sea más prominente, necesitas darle estilo.

Selecciona el elemento `fieldset legend`. Establece su `font-size` en `1.3rem` y su `font-weight` en `600`.

`font-weight` controla el grosor del texto. 

Los valores comunes de `font-weight` son:

- `normal` → Peso regular (predeterminado)
- `bold` → Texto en negrita
- `lighter` → Más claro que el elemento padre
- `bolder` → Más negrita que el elemento padre
- Numeric values like:
  - `100` (delgado)
  - `400` (normal)
  - `700` (negrita)
  - `900` (extra negrita)

# --hints--

Debes tener un selector `fieldset legend`.

```js
assert(new __helpers.CSSHelp(document).getStyle('fieldset legend'));
```

Establece la propiedad `font-size` en `1.3rem`.

```js
const style = new __helpers.CSSHelp(document).getStyle('fieldset legend');
assert.strictEqual(style?.fontSize, '1.3rem');
```

Establece la propiedad `font-weight` en `600`.

```js
const style = new __helpers.CSSHelp(document).getStyle('fieldset legend');
assert.strictEqual(style?.fontWeight, '600');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Information: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```

```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

.container {
    background-color: #ffffff1a;
    width: 80%;
    max-width: 600px;
    border-radius: 10px;
    margin: 20px auto;
    padding: 10px 20px;
    box-shadow: 0 5px 15px black;  
}

.center {
  text-align: center;
}

.description {
  font-size: 1.2rem;
}

fieldset {
  border: 1px solid gray;
  border-radius: 5px;
  margin: 20px 0;
  padding: 20px;
}

--fcc-editable-region--


--fcc-editable-region--
```
