---
id: 66e2df8c5676fe51cead8ec8
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Ora è il momento di testare la tua funzione `getVowelCount`.

Crea una variabile `vowelCount` e assegnale il risultato della chiamata alla funzione `getVowelCount` con l'argomento `"Apples are tasty fruits"`

Dopodiché, stampa nella console il seguente valore: `"Vowel Count: [vowel count goes here]"`. Sostituisci `[vowel count goes here]` con il nome reale della variabile. Puoi scegliere di usare template stringhe o la concatenazione di stringhe con l'operatore `+`.

# --hints--

Dovresti creare una variabile `vowelCount`.

```js
assert.isNotNull(vowelCount)
```

La tua variabile `vowelCount` dovrebbe essere impostata al risultato di `getVowelCount("Apples are tasty fruits")`.

```js
assert.equal(vowelCount, getVowelCount("Apples are tasty fruits"));
```

Dovresti stampare nella console la variabile `vowelCount` con il parametro `"Vowel Count: ${vowelCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Vowel\s+Count:\s+('|"|`)\s+\+\s+vowelCount|`Vowel\s+Count:\s+\${vowelCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
