---
id: 66e2eab8a5638f57b637b7cc
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Ora dovresti contare il numero di punteggiature.

Crea una funzione `getPunctuationCount` con un parametro `sentence`.

All'interno della funzione, crea un loop per contare il numero di caratteri di punteggiatura nel `sentence` che verrà passato alla funzione quando viene chiamata. Ai fini di questo esercizio, un carattere di punteggiatura è qualsiasi carattere che non sia uno spazio (`" "`) o una lettera.

La tua funzione `getPunctuationCount` deve restituire un numero.

# --hints--

Dovresti creare una funzione `getPunctuationCount`.

```js
assert.isFunction(getPunctuationCount);
```

La tua funzione `getPunctuationCount` dovrebbe avere un parametro `sentence`.

```js
assert.match(getPunctuationCount.toString(), /sentence/);
```

La tua funzione `getPunctuationCount` dovrebbe restituire un numero.

```js
assert.isNumber(getPunctuationCount("Coding is fun!"))
```

Quando la frase è `"What's going on here?"`, la funzione `getPunctuationCount` dovrebbe restituire `2`.

```js
assert.strictEqual(getPunctuationCount("What's going on here?"), 2);
```

Quando la frase è `"What????!"`, la funzione `getPunctuationCount` dovrebbe restituire `5`.

```js
assert.strictEqual(getPunctuationCount("What????!"), 5);
```

La tua funzione `getPunctuationCount` dovrebbe restituire il conteggio corretto della punteggiatura per qualsiasi frase.

```js
assert.strictEqual(getPunctuationCount("Be quick, sign up! freeCodeCamp awaits, friend!!!"), 6);
assert.strictEqual(getPunctuationCount("Guess what? freeCodeCamp is launching a new cert soon!"), 2);
assert.strictEqual(getPunctuationCount("freeCodeCamp, again? It's incredible!"), 4);
assert.strictEqual(getPunctuationCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
