---
id: 66e2e99caa0cb3570f4822b4
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Ora è il momento di testare la tua funzione `getConsonantCount`.

Crea una variabile `consonantCount` e assegnale il risultato della chiamata della funzione `getConsonantCount` con l'argomento `"Coding is fun"`

Dopodiché, stampa sulla console il seguente valore: `"Consonant Count: [Consonant count goes here]"`. Sostituisci `[Consonant count goes here]` con il nome reale della variabile. Puoi scegliere di usare stringhe template o la concatenazione di stringhe con l'operatore `+`.

# --hints--

Dovresti creare una variabile `consonantCount`.

```js
assert.isNotNull(consonantCount)
```

La tua variabile `consonantCount` dovrebbe essere impostata al risultato di `getConsonantCount("Coding is fun")`.

```js
assert.equal(consonantCount, getConsonantCount("Coding is fun"));
```

Dovresti stampare sulla console la variabile `consonantCount` con il parametro `"Consonant Count: ${consonantCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Consonant\s+Count:\s+('|"|`)\s+\+\s+consonantCount|`Consonant\s+Count:\s+\${consonantCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
