---
id: 66e2edb23013ca596f8d9568
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Ora è il momento di testare la tua funzione `getPunctuationCount`.

Crea una variabile `punctuationCount` e assegnale il risultato della chiamata della funzione `getPunctuationCount` con l'argomento `"WHAT?!?!?!?!?"`

Dopodiché, stampa nella console il seguente valore: `"Punctuation Count: [Punctuation count goes here]"`. Sostituisci `[Punctuation count goes here]` con il nome reale della variabile. Puoi scegliere di usare template string o la concatenazione di stringhe con l'operatore `+`.

# --hints--

Dovresti creare una variabile `punctuationCount`.

```js
assert.isNotNull(punctuationCount)
```

La tua variabile `punctuationCount` dovrebbe essere impostata al risultato di `getPunctuationCount("WHAT?!?!?!?!?")`.

```js
assert.equal(punctuationCount, getPunctuationCount("WHAT?!?!?!?!?"));
```

Dovresti stampare nella console il seguente valore: `"Punctuation Count: [Punctuation count goes here]"`. Sostituisci `[Punctuation count goes here]` con il nome reale della variabile. Assicurati di usare la sintassi corretta per la concatenazione di stringhe.

```js
assert.match(code, /console\.log\((?:('|"|`)Punctuation\s+Count:\s+('|"|`)\s+\+\s+punctuationCount|`Punctuation\s+Count:\s+\${punctuationCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
