---
id: 66e2f376df6f315ee81de81a
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Metti alla prova il tuo `getWordCount` creando una variabile `wordCount` impostata sulla chiamata della funzione `getWordCount` con la frase `"I love freeCodeCamp"`.

Dopodiché, stampa sulla console il seguente valore: `"Word Count: [Word count goes here]"`. Sostituisci `[Word count goes here]` con il nome reale della variabile. Puoi scegliere di usare stringhe template o la concatenazione di stringhe con l’operatore `+`.

Con questo, il tuo progetto di analizzatore di frasi è completato!

# --hints--

Dovresti creare una variabile `wordCount`.

```js
assert.isNotNull(wordCount)
```

La tua variabile `wordCount` dovrebbe essere impostata sul risultato di `getWordCount("I love freeCodeCamp")`.

```js
assert.equal(wordCount, getWordCount("I love freeCodeCamp"));
```

Dovresti stampare sulla console la variabile `wordCount` con il parametro `"Word Count: ${wordCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Word\s+Count:\s+('|"|`)\s+\+\s+wordCount|`Word\s+Count:\s+\${wordCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

function getWordCount(sentence) {
  if (sentence.trim() === '') {
    return 0;
  }
  
  const words = sentence.trim().split(' ');
  let count = 0;

  for (const word of words) {
    if (word !== '') {
      count++;
    }
  }

  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

function getWordCount(sentence) {
  if (sentence.trim() === '') {
    return 0;
  }
  
  const words = sentence.trim().split(' ');
  let count = 0;

  for (const word of words) {
    if (word !== '') {
      count++;
    }
  }

  return count;
}

const wordCount = getWordCount("I love freeCodeCamp");
console.log(`Word Count: ${wordCount}`);
```
