---
id: 66e2e5d890ac8a54c9d703bd
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

È il momento di contare le consonanti. Crea una `getConsonantCount` funzione con un parametro `sentence`.

All'interno della funzione, usa un loop per contare il numero di consonanti nel `sentence` che verrà passato alla funzione quando viene chiamata. Una consonante è qualsiasi lettera che non sia uno dei seguenti caratteri: `"aeiou"`.

La tua funzione `getConsonantCount` deve restituire un numero.

# --hints--

Dovresti creare una funzione `getConsonantCount`.

```js
assert.isFunction(getConsonantCount);
```

La tua funzione `getConsonantCount` dovrebbe avere un parametro `sentence`.

```js
assert.match(getConsonantCount.toString(), /sentence/);
```

La tua funzione `getConsonantCount` dovrebbe restituire un numero.

```js
assert.isNumber(getConsonantCount("Coding is fun"))
```

Quando la frase è `"Coding is fun"`, la funzione `getConsonantCount` dovrebbe restituire `7`.

```js
assert.strictEqual(getConsonantCount("Coding is fun"), 7);
```

Quando la frase è `"Hello, World!"`, la funzione `getConsonantCount` dovrebbe restituire `7`.

```js
assert.strictEqual(getConsonantCount("Hello, World!"), 7);
```

Il conteggio delle consonanti deve essere case-insensitive.

```js
assert.strictEqual(getConsonantCount("Apples are tasty fruits"), 13);
assert.strictEqual(getConsonantCount("freeCodeCamp is awesome"), 11);
```

La tua funzione `getConsonantCount` dovrebbe restituire il conteggio corretto delle consonanti per qualsiasi frase.

```js
assert.strictEqual(getConsonantCount("I went to the store"), 9);
assert.strictEqual(getConsonantCount("The quick brown fox jumps over the lazy dog"), 24);
assert.strictEqual(getConsonantCount("The cat in the hat"), 9);
assert.strictEqual(getConsonantCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
