---
id: 66e2d680e129e1423116a541
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop costruirai un analizzatore di frasi che prenderà una frase e calcolerà il numero di parole, vocali, consonanti e segni di punteggiatura.

Per cominciare, crea una funzione `getVowelCount` con un parametro chiamato `sentence`. La tua funzione dovrebbe restituire il numero totale di vocali in una frase.

# --hints--

Dovresti creare una funzione `getVowelCount`.

```js
assert.isFunction(getVowelCount);
```

La tua funzione `getVowelCount` dovrebbe avere un parametro `sentence`.

```js
assert.match(getVowelCount.toString(), /sentence/);
```

La tua funzione `getVowelCount` dovrebbe restituire un numero.

```js
assert.isNumber(getVowelCount("Apples are tasty fruits"))
```

Quando la frase è `"Apples are tasty fruits"`, la funzione `getVowelCount` dovrebbe restituire `7`.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
```

Quando la frase è `"Hello, World!"`, la funzione dovrebbe restituire `3`.

```js
assert.strictEqual(getVowelCount("Hello, World!"), 3);
```

Il conteggio delle vocali dovrebbe essere case-insensitive.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
assert.strictEqual(getVowelCount("freeCodeCamp is awesome"), 10);
```

La tua funzione `getVowelCount` dovrebbe restituire il conteggio corretto delle vocali per qualsiasi frase.

```js
assert.strictEqual(getVowelCount("I went to the store"), 6);
assert.strictEqual(getVowelCount("The quick brown fox jumps over the lazy dog"), 11);
assert.strictEqual(getVowelCount("The cat in the hat"), 5);
assert.strictEqual(getVowelCount(""), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
