---
id: 66e2f06b191e305c00574e4d
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Infine, conta il numero di parole creando una funzione `getWordCount` con un parametro `sentence`. La funzione deve restituire il numero totale di parole nella frase passata quando viene chiamata.

# --hints--

Dovresti creare una funzione `getWordCount`.

```js
assert.isFunction(getWordCount);
```

La tua funzione `getWordCount` dovrebbe avere un parametro `sentence`.

```js
assert.match(getWordCount.toString(), /sentence/);
```

La tua funzione `getWordCount` dovrebbe restituire un numero.

```js
assert.isNumber(getWordCount("I love freeCodeCamp"))
```

Quando la frase è `"When are you gonna start learning to code?"`, la funzione `getWordCount` dovrebbe restituire `8`.

```js
assert.strictEqual(getWordCount("When are you gonna start learning to code?"), 8);
```

Quando la frase è `"What's going on?"`, la funzione `getWordCount` dovrebbe restituire `3`.

```js
assert.strictEqual(getWordCount("What's going on?"), 3);
```

Il conteggio delle parole deve essere insensibile alle maiuscole.

```js
assert.strictEqual(getWordCount("freeCodeCamp offers free coding tutorials online"), 6);
assert.strictEqual(getWordCount("You can learn HTML, CSS, JavaScript, and more on freeCodeCamp"), 10);
```

La tua funzione `getWordCount` dovrebbe restituire il conteggio corretto delle parole per qualsiasi frase.

```js
assert.strictEqual(getWordCount("freeCodeCamp has a great community of kind people"), 8);
assert.strictEqual(getWordCount("The freeCodeCamp curriculum is constantly updated"), 6);
assert.strictEqual(getWordCount("freeCodeCamp teaches both front-end and back-end development"), 7);
```

La tua funzione `getWordCount` dovrebbe restituire il conteggio corretto delle parole per una stringa vuota o una stringa contenente solo spazi.

```js
assert.strictEqual(getWordCount(""), 0);
assert.strictEqual(getWordCount("   "), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
