---
id: 6153893900438b4643590367
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Вирівняйте свій `.header` текст по центру. Зробіть текст у верхньому регістрі, використовуючи властивість `text-transform` зі значенням `uppercase`.

Додайте відступ `32px` з усіх боків. Встановіть фон `#0a0a23`, а колір тексту — `#fff`.

Додайте `border-bottom` зі значенням `4px solid #fdb347`.

# --hints--

У вас має бути селектор `.header`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

Ваш селектор `.header` має мати властивість `text-align`, встановлену у значення `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign, 'center');
```

Ваш селектор `.header` має мати властивість `text-transform`, встановлену у значення `uppercase`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform, 'uppercase');
```

Ваш селектор `.header` має мати властивість `padding`, встановлену у значення `32px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.padding, '32px');
```

Ваш селектор `.header` має мати властивість `background-color`, встановлену у значення `#0a0a23`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor, 'rgb(10, 10, 35)');
```

Ваш селектор `.header` має мати властивість `color`, встановлену у значення `#fff`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.color, 'rgb(255, 255, 255)');
```

Ваш селектор `.header` має мати властивість `border-bottom`, встановлену у значення `4px solid #fdb347`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom, '4px solid rgb(253, 179, 71)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}
body  {
  margin: 0;
  font-family: sans-serif;
  background-color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
