---
id: 61537bb9b1a29430ac15ad38
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

У вашому елементі `head` додайте елемент `title` з текстом, встановленим у `Photo Gallery`, а також елемент `link`, щоб додати ваш файл `styles.css` на сторінку.

# --hints--

Ваш елемент `link` має мати атрибут `href` зі значенням `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

Ваш код має містити елемент `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Ваш проєкт має мати назву `Photo Gallery`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'photo gallery')
```

Пам’ятайте, що регістр і правопис мають значення для назви.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Photo Gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
