---
id: 61537a8054753e2f1f2a1574
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

У вашому елементі `head` додайте тег `meta` з атрибутом `name`, встановленим у `viewport`, та атрибутом `content`, встановленим у `width=device-width, initial-scale=1`.

Також додайте тег `meta` з атрибутом `charset`, встановленим у `UTF-8`.

# --hints--

У вас має бути два елементи `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Один елемент `meta` повинен мати атрибут `name`, встановлений у `viewport`, і атрибут `content`, встановлений у `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content').match(/width=device-width,\s?initial-scale=1(.0)?/) && !m?.getAttribute('charset'));
assert.exists(target);
```

Інший елемент `meta` повинен мати атрибут `charset`, встановлений у `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
