---
id: 61537485c4f2a624f18d7794
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Почніть зі свого стандартного шаблонного коду HTML. Додайте оголошення `DOCTYPE`, елемент `html`, елемент `head` та елемент `body`.

Додайте атрибут `lang` до початкового тегу `<html>` зі значенням `en`.

# --hints--

Ваш код має містити референс `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Після референса `DOCTYPE` має бути пробіл.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Визначте `html` як тип документа.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Закрийте оголошення `DOCTYPE` кінцевим тегом `>` після типу.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Ваш елемент `html` має мати початковий тег.

```js
assert(code.match(/<html\s*>|<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ваш елемент `html` має мати кінцевий тег.

```js
assert(code.match(/<\/html\s*>/));
```

Початковий тег `<html>` має містити атрибут `lang` зі значенням `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/));
```

Оголошення `DOCTYPE` має бути на початку вашого HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Ви маєте мати початковий тег `head`.

```js
assert(code.match(/<head\s*>/i));
```

Ви маєте мати кінцевий тег `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Ви маєте мати початковий тег `body`.

```js
assert(code.match(/<body\s*>/i));
```

Ви маєте мати кінцевий тег `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Елементи `head` та `body` мають бути братськими.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Елемент `head` має бути всередині елемента `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Елемент `body` має бути всередині елемента `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
