---
id: 61537c5f81f0cf325b4a854c
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Додайте елемент `header` всередину елемента `body` і призначте йому клас `header`.

Всередині `header` створіть `h1` з `css flexbox photo gallery` як текстом.

# --hints--

У вашому елементі `header` має бути елемент `body`.

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

Ваш елемент `header` має мати `class` зі значенням `header`.

```js
assert(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

Ваш елемент `header` має містити всередині елемент `h1`.

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

У вашому елементі `h1` має бути текст `css flexbox photo gallery`.

```js
assert(document?.querySelector('.header')?.querySelector('h1')?.innerText === 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
