---
id: 56533eb9ac21ba0edf2244b3
title: 摂氏から華氏への変換器を作る
challengeType: 26
dashedName: build-a-celsius-to-fahrenheit-converter
---

# --description--

このラボでは、摂氏から華氏に温度を変換する関数を書きます。摂氏から華氏への変換式は次の通りです：

```js
fahrenheit = celsius * (9/5) + 32
```

**目的：** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー：**

1. `convertCtoF` という名前の関数を作成してください。
2. `convertCtoF` は摂氏の温度を表す1つの数値の引数を受け取る必要があります。
3. `convertCtoF` は数値を返す必要があります。

# --hints--

`convertCtoF` という名前の関数を作成してください。

```js
assert.isFunction(convertCtoF);
```

`convertCtoF` は1つのパラメータを受け取る必要があります。

```js
assert.lengthOf(convertCtoF, 1);
```

`convertCtoF(0)` は数値を返す必要があります。

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` は `-22` の値を返す必要があります。

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` は `14` の値を返す必要があります。

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` は `32` の値を返す必要があります。

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` は `68` の値を返す必要があります。

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` は `86` の値を返す必要があります。

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
