---
id: 68da6138d70960c197ec5eed
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

接著，你需要驗證配對數量至少為 1，因為你無法用零或負數的配對來產生小括弧組合。

為你的程式碼添加另一個 `if` 述語來檢查 `pairs` 是否小於 `1`。如果此條件為真，傳回字串 `The number of pairs should be at least 1`。

# --hints--

你應該在你的 `gen_parentheses` 函式中有第二個 `if` 述語。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[1]
`) })
```

你的第二個 `if` 述語應該檢查 `pairs` 是否小於 `1`。

```js
({ test: () => runPython(`
cond = _Node(_code).find_function("gen_parentheses").find_ifs()[1].find_conditions()[0]
assert cond.is_equivalent("pairs < 1") or cond.is_equivalent("1 > pairs")
`) })
```

當 `pairs` 小於 `1` 時，你應該傳回字串 `The number of pairs should be at least 1`。

```js
({ test: () => runPython(`
msg = 'The number of pairs should be at least 1'
assert gen_parentheses(0) == msg and gen_parentheses(-1) == msg
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    
    return []
--fcc-editable-region--
```
