---
id: 68da6138d70960c197ec5ef0
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

現在你將實作主要的 BFS 迴圈。創建一個 `while` 迴圈，只要 `queue` 不為空（當 `queue` 求值為 `True`）就持續執行。

在迴圈中，列印 `queue`。

# --hints--

你應該在你的 `gen_parentheses` 函式中有一個 `while` 迴圈。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0]
`) })
```

你的 `while` 迴圈應該有條件 `queue`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_conditions()[0].is_equivalent("queue")
`) })
```

你應該在你的 `while` 迴圈中列印 `queue`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_call("print(queue)")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    queue = [('', 0, 0)]
    result = []
    
    
    return result
--fcc-editable-region--
```
