---
id: 68da6138d70960c197ec5eec
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

在實作核心演算法之前，你需要驗證輸入。`pairs` 參數應該是整數，因為它表述要產生的小括弧對組數量。

在你的函式（程式）開頭添加一個 `if` 述語，以檢查 `pairs` 是否不是整數。為此使用 `isinstance()` 函式（程式）。

如果條件為真，傳回字串 `The number of pairs should be an integer`。

# --hints--

你應該在你的 `gen_parentheses` 函式中有一個 `if` 述語。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0]
`) })
```

你的 `if` 述語應該使用 `not isinstance(pairs, int)` 來檢查 `pairs` 是否不是整數。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(pairs, int)")
`) })
```

當輸入不是整數時，你應該傳回字串 `The number of pairs should be an integer`。

```js
({ test: () => runPython(`
msg = 'The number of pairs should be an integer'
args = [1.1, "1", [1, 2], {}]
assert  all(gen_parentheses(arg) == msg for arg in args)
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    
    return []
--fcc-editable-region--
```
