---
id: 68da6138d70960c197ec5ef2
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

如果目前的字串尚未完成，你需要探索下一個可能的狀態。為你的 `if` 述語添加一個 `else` 語句。

在 `else` 區塊中，你將處理添加開啟小括弧。為你的程式碼添加一個 `if` 述語來檢查 `opens_used < pairs`。這可確保你不會使用超過允許數量的開啟小括弧。

如果此條件為真，將新的元組 `(current + '(', opens_used + 1, closes_used)` 附加到 `queue`。這表述了添加開括號後的狀態。

# --hints--

你應該在 `while` 迴圈內的 `if` 述語中有一個 `else` 語句。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse
`) })
```

你應該在 `else` 區塊內有一個 `if` 述語，檢查 `opens_used < pairs`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1].is_equivalent("opens_used < pairs")
`) })
```

當條件為真時，你應該將 `(current + '(', opens_used + 1, closes_used)` 附加到佇列中。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[1].has_call("queue.append((current + '(', opens_used + 1, closes_used))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        if len(current) == 2 * pairs:
            result.append(current)
        
--fcc-editable-region--
    return result

print(gen_parentheses(1))
```
