---
id: 68da6138d70960c197ec5eef
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

對於廣度優先搜尋方法，你將使用佇列來追蹤在建置小括弧組合時的不同狀態。每個狀態將表述為包含三個元素的元組：

- 目前正在建立的字串
- 目前使用的開啟小括號數量
- 目前使用的結束小括弧數量

創建一個名為 `queue` 的變數，並以包含一個元組：`('', 0, 0)` 的列表初始化它。這表述了起始狀態，字串為空且使用了零個小括弧。

# --hints--

你應該在你的 `gen_parentheses` 函式中宣告一個名為 `queue` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("queue")
`) })
```

你應該用包含元組 `('', 0, 0)` 的列表來初始化 `queue`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("queue").is_equivalent("queue = [('', 0, 0)]")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    result = []
    return result
--fcc-editable-region--
```
