---
id: 68da6138d70960c197ec5eee
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

現在你將設定資料結構來儲存你的結果。創建一個名為 `result` 的變數，並將其初始化為空列表。這個列表將儲存你產生的所有有效小括弧組合。

更新你的傳回述語以傳回 `result`，而非空的列表。

# --hints--

你應該在你的 `gen_parentheses` 函式中宣告一個名為 `result` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("result")
`) })
```

你應該將 `result` 初始化為空的列表。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("result").is_equivalent("result = []")
`) })
```

你應該在你的函式結尾傳回 `result`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_return("result")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    return []
--fcc-editable-region--
```
