---
id: 68da6138d70960c197ec5ef3
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

現在你需要控制代碼來添加結束小括弧。關鍵規則是你只能在維護平衡的情況下添加結束小括弧，這表示使用的小括弧數量必須多於結束小括弧。

仍在 `else` 區塊中，添加第二個 `if` 述語以檢查 `closes_used < opens_used`。

如果此條件為真，則將另一個新的元組附加到 `queue`：`(current + ')'，opens_used，closes_used + 1)`。這表述添加閉括號後的狀態。

# --hints--

你應該在你的 `else` 區塊內有第二個 `if` 述語。

```js
({ test: () => runPython(`
import ast
assert isinstance(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1], ast.If)
`) })
```

你的第二個巢狀 `if` 述語應該檢查 `closes_used < opens_used`。

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1].test).is_equivalent("closes_used < opens_used")
`) })
```

當條件為真時，你應該將 `(current + ')', opens_used, closes_used + 1)` 附加到佇列。

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1]).has_call("queue.append((current + ')', opens_used, closes_used + 1))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        
        if len(current) == 2 * pairs:
            result.append(current)
--fcc-editable-region--
        else:
            if opens_used < pairs:
                queue.append((current + '(', opens_used + 1, closes_used))
            
--fcc-editable-region--
    return result

print(gen_parentheses(1))
```
