---
id: 68da6138d70960c197ec5ef1
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

現在你需要判斷何時已經建立完整的小括弧組合。完整的組合長度相等於配對數量的兩倍（因為每個配對包含一個開啟和一個關閉的小括弧）。

在你的 `while` 迴圈中，為你的 `if` 述語添加檢查 `len(current) == 2 * pairs` 的條件。當此條件為真時，將 `current` 附加到 `result` 列表。

# --hints--

你應該在你的 `while` 迴圈中放置一個 `if` 述語。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0]
`) })
```

你的 `if` 述語應該檢查 `len(current) == 2 * pairs` 是否成立。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("len(current) == 2 * pairs")
`) })
```

你應該在新的 `if` 述語內將 `current` 附加到 `result` 列表中。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("result.append(current)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        
--fcc-editable-region--    
    return result

print(gen_parentheses(1))
```
