---
id: 68da604b8b6012a33e2344fb
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個工作坊中，你將實作一個函式，使用廣度優先搜尋（BFS）方法來產生所有有效的小括弧組合。例如，兩對小括弧的有效組合是 `(())` 和 `()()`。

首先建立一個名為 `gen_parentheses` 的函式，並帶有一個參數 `pairs`。目前，從函式中傳回一個空的列表。

# --hints--

你應該定義一個名為 `gen_parentheses` 的函式（程式）。

```js
({ test: () => runPython(`
assert _Node(_code).has_function("gen_parentheses")
`) })
```

你的 `gen_parentheses` 函式（程式）應該有一個名為 `pairs` 的單一參數。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_args("pairs")
`) })
```

你的函式（程式）應該傳回一個空的列表。

```js
({ test: () => runPython(`
foo = _Node(_code).find_function("gen_parentheses")
assert foo.has_return("[]") or foo.has_return("list()")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
