---
id: 68e00f88cf82083df8e0497b
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

在你繼續進行 BFS 邏輯之前，呼叫 `gen_parentheses(1)` 並將結果列印到主控台。

# --hints--

你應該將 `gen_parentheses(1)` 列印到主控台。

```js
({ test: () => runPython(`
assert _Node(_code).has_call('print(gen_parentheses(1))')
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []

    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)

    return result
--fcc-editable-region--

--fcc-editable-region--
```
