---
id: 68e00c5349cbddfef6dbc053
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

在你的 `while` 迴圈中，使用 `queue.pop(0)` 來移除並取得佇列中的第一個元素。這實作了 BFS 特有的先進先出（FIFO）行為。

將此 tuple 解包為三個變數：`current`、`opens_used` 和 `closes_used`。

# --hints--

你應該使用 `queue.pop(0)` 來取得佇列中的第一個元素。

```js
({ test: () => runPython(`
assert "queue.pop(0)" in str(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0])
`) })
```

你應該將結果解包到三個變數：`current`、`opens_used` 和 `closes_used`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_stmt("current, opens_used, closes_used = queue.pop(0)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        
--fcc-editable-region--
    return result
```
