---
id: 56104e9e514f539506016a5c
title: Iterare numeri dispari con un for loop
challengeType: 1
forumTopicId: 18212
dashedName: iterate-odd-numbers-with-a-for-loop
---

# --description--

I for loop non devono necessariamente iterare uno alla volta. Modificando il nostro `final-expression`, possiamo contare con numeri pari.

Inizieremo da `i = 0` e itereremo finché `i < 10`. Incrementeremo `i` di 2 a ogni iterazione con `i += 2`.

```js
const ourArray = [];

for (let i = 0; i < 10; i += 2) {
  ourArray.push(i);
}
```

`ourArray` conterrà ora `[0, 2, 4, 6, 8]`. Cambiamo il nostro `initialization` così possiamo contare con numeri dispari.

# --instructions--

Inserisci i numeri dispari da 1 a 9 in `myArray` usando un `for` loop.

# --hints--

Dovresti usare un `for` loop per questo.

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

`myArray` dovrebbe essere uguale a `[1, 3, 5, 7, 9]`.

```js
assert.deepEqual(myArray, [1, 3, 5, 7, 9]);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 1; i < 10; i += 2) {
  myArray.push(i);
}
```
