---
id: 56105e7b514f539506016a5e
title: Contare all'indietro con un for loop
challengeType: 1
forumTopicId: 16808
dashedName: count-backwards-with-a-for-loop
---

# --description--

Un for loop può anche contare all'indietro, purché possiamo definire le condizioni corrette.

Per decrementare di due a ogni iterazione, dovremo modificare la nostra inizializzazione, condizione ed espressione finale.

Inizieremo da `i = 10` e itereremo finché `i > 0`. Decrementeremo `i` di 2 a ogni loop con `i -= 2`.

```js
const ourArray = [];

for (let i = 10; i > 0; i -= 2) {
  ourArray.push(i);
}
```

`ourArray` conterrà ora `[10, 8, 6, 4, 2]`. Cambiamo la nostra inizializzazione e l'espressione finale così da poter contare all'indietro di due in due per creare un array di numeri dispari in ordine decrescente.

# --instructions--

Inserisci i numeri dispari da 9 a 1 in `myArray` usando un `for` loop.

# --hints--

Dovresti usare un `for` loop per questo.

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

Dovresti usare il metodo array `push`.

```js
assert(__helpers.removeJSComments(code).match(/myArray.push/));
```

`myArray` dovrebbe essere uguale a `[9, 7, 5, 3, 1]`.

```js
assert.deepEqual(myArray, [9, 7, 5, 3, 1]);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 9; i > 0; i -= 2) {
  myArray.push(i);
}
```
