---
id: 56533eb9ac21ba0edf2244d9
title: "Confronti con l'operatore logico or"
challengeType: 1
forumTopicId: 16800
dashedName: comparisons-with-the-logical-or-operator
---

# --description--

L'<dfn>operatore logico or</dfn> (`||`) restituisce `true` se almeno uno dei <dfn>operandi</dfn> è `true`. Altrimenti, restituisce `false`.

L'<dfn>operatore logico or</dfn> è composto da due simboli pipe: (`||`). Di solito si trova tra i tasti Backspace e Invio.

Il pattern qui sotto dovrebbe esserti familiare dai waypoint precedenti.

```js
if (num > 10) {
  return "No";
}
if (num < 5) {
  return "No";
}
return "Yes";
```

Questo codice restituirà `Yes` se `num` è compreso tra `5` e `10` (inclusi `5` e `10`). La stessa logica può essere scritta con l'<dfn>operatore logico or</dfn>.

```js
if (num > 10 || num < 5) {
  return "No";
}
return "Yes";
```

# --instructions--

Combina le due istruzioni `if` in una sola istruzione che restituisce la stringa `Outside` se `val` non è compreso tra `10` e `20`, inclusi. Altrimenti, restituisci la stringa `Inside`.

# --hints--

Dovresti usare l'operatore `||` una volta

```js
assert(__helpers.removeJSComments(code).match(/\|\|/g).length === 1);
```

Dovresti avere solo una istruzione `if`

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`testLogicalOr(0)` dovrebbe restituire la stringa `Outside`

```js
assert(testLogicalOr(0) === 'Outside');
```

`testLogicalOr(9)` dovrebbe restituire la stringa `Outside`

```js
assert(testLogicalOr(9) === 'Outside');
```

`testLogicalOr(10)` dovrebbe restituire la stringa `Inside`

```js
assert(testLogicalOr(10) === 'Inside');
```

`testLogicalOr(15)` dovrebbe restituire la stringa `Inside`

```js
assert(testLogicalOr(15) === 'Inside');
```

`testLogicalOr(19)` dovrebbe restituire la stringa `Inside`

```js
assert(testLogicalOr(19) === 'Inside');
```

`testLogicalOr(20)` dovrebbe restituire la stringa `Inside`

```js
assert(testLogicalOr(20) === 'Inside');
```

`testLogicalOr(21)` dovrebbe restituire la stringa `Outside`

```js
assert(testLogicalOr(21) === 'Outside');
```

`testLogicalOr(25)` dovrebbe restituire la stringa `Outside`

```js
assert(testLogicalOr(25) === 'Outside');
```

# --seed--

## --seed-contents--

```js
function testLogicalOr(val) {
  // Only change code below this line

  if (val) {
    return "Outside";
  }

  if (val) {
    return "Outside";
  }

  // Only change code above this line
  return "Inside";
}

testLogicalOr(15);
```

# --solutions--

```js
function testLogicalOr(val) {
  if (val < 10 || val > 20) {
    return "Outside";
  }
  return "Inside";
}
```
