---
id: 56533eb9ac21ba0edf2244b5
title: Escape delle citazioni letterali nelle stringhe
challengeType: 1
forumTopicId: 17568
dashedName: escaping-literal-quotes-in-strings
---

# --description--

Quando definisci una stringa devi iniziare e finire con una citazione singola o doppia. Cosa succede quando ti serve una citazione letterale: `"` o `'` all’interno della tua stringa?

In JavaScript, puoi <dfn>fare escape</dfn> di una citazione per evitare che venga interpretata come la fine della stringa mettendo una <dfn>barra rovesciata</dfn> (`\`) davanti alla citazione.

```js
const sampleStr = "Alan said, \"Peter is learning JavaScript\".";
```

Questo indica a JavaScript che la citazione successiva non è la fine della stringa, ma deve invece apparire all’interno della stringa. Quindi, se la stampassi sulla console, otterresti:

```js
Alan said, "Peter is learning JavaScript".
```

# --instructions--

Usa le <dfn>barre rovesciate</dfn> per assegnare una stringa alla variabile `myStr` in modo che, se la stampassi sulla console, vedresti:

```js
I am a "double quoted" string inside "double quotes".
```

# --hints--

Dovresti usare due virgolette doppie (`"`) e quattro virgolette doppie con escape (`\"`).

```js
assert(__helpers.removeJSComments(code).match(/\\"/g).length === 4 && __helpers.removeJSComments(code).match(/[^\\]"/g).length === 2);
```

La variabile `myStr` dovrebbe contenere la stringa: `I am a "double quoted" string inside "double quotes".`

```js
assert(/I am a "double quoted" string inside "double quotes(\."|"\.)$/.test(myStr));
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "I am a \"double quoted\" string inside \"double quotes\".";
```
