---
id: 56533eb9ac21ba0edf2244b6
title: Sequenze di escape nelle stringhe
challengeType: 1
forumTopicId: 17567
dashedName: escape-sequences-in-strings
---

# --description--

Le virgolette non sono gli unici caratteri che possono essere <dfn>escapati</dfn> all'interno di una stringa. Le sequenze di escape ti permettono di usare caratteri che altrimenti non potresti inserire in una stringa.

<table><thead><tr><th>Codice</th><th>Output</th></tr></thead><tbody><tr><td><code>\'</code></td><td>virgolette singole</td></tr><tr><td><code>\"</code></td><td>virgolette doppie</td></tr><tr><td><code>\\</code></td><td>barra rovesciata</td></tr><tr><td><code>\n</code></td><td>nuova riga</td></tr><tr><td><code>\t</code></td><td>tabulazione</td></tr><tr><td><code>\r</code></td><td>ritorno a capo</td></tr><tr><td><code>\b</code></td><td>backspace</td></tr><tr><td><code>\f</code></td><td>form feed</td></tr></tbody></table>

*Nota che la barra rovesciata stessa deve essere escapata per poter essere mostrata come barra rovesciata.*

# --instructions--

Assegna le seguenti tre righe di testo nella singola variabile `myStr` usando le sequenze di escape.

<pre>
FirstLine
    \SecondLine
ThirdLine
</pre>

Dovrai usare le sequenze di escape per inserire correttamente i caratteri speciali. Dovrai anche rispettare gli spazi come mostrato sopra, senza spazi tra le sequenze di escape o le parole.

**Nota:** L'indentazione per `SecondLine` si ottiene con il carattere di escape tabulazione, non con spazi.

# --hints--

`myStr` non dovrebbe contenere spazi

```js
assert(!/ /.test(myStr));
```

`myStr` dovrebbe contenere le stringhe `FirstLine`, `SecondLine` e `ThirdLine` (ricorda la distinzione tra maiuscole e minuscole)

```js
assert(
  /FirstLine/.test(myStr) && /SecondLine/.test(myStr) && /ThirdLine/.test(myStr)
);
```

`FirstLine` dovrebbe essere seguito dal carattere di nuova riga `\n`

```js
assert(/FirstLine\n/.test(myStr));
```

`myStr` dovrebbe contenere un carattere tabulazione `\t` che segue un carattere di nuova riga

```js
assert(/\n\t/.test(myStr));
```

`SecondLine` dovrebbe essere preceduto dal carattere barra rovesciata `\`

```js
assert(/\\SecondLine/.test(myStr));
```

Dovrebbe esserci un carattere di nuova riga tra `SecondLine` e `ThirdLine`

```js
assert(/SecondLine\nThirdLine/.test(myStr));
```

`myStr` dovrebbe contenere solo i caratteri mostrati nelle istruzioni

```js
assert(myStr === 'FirstLine\n\t\\SecondLine\nThirdLine');
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "FirstLine\n\t\\SecondLine\nThirdLine";
```
