---
id: 56533eb9ac21ba0edf2244c7
title: Consultare le proprietà di un oggetto con la notazione a punto
challengeType: 1
forumTopicId: 16164
dashedName: accessing-object-properties-with-dot-notation
---

# --description--

Ci sono due modi per consultare le proprietà di un oggetto: la notazione a punto (`.`) e la notazione a parentesi (`[]`), simile a un array.

La notazione a punto si usa quando conosci in anticipo il nome della proprietà che vuoi consultare.

Ecco un esempio di come usare la notazione a punto (`.`) per leggere una proprietà di un oggetto:

```js
const myObj = {
  prop1: "val1",
  prop2: "val2"
};

const prop1val = myObj.prop1;
const prop2val = myObj.prop2;
```

`prop1val` avrebbe come valore la stringa `val1`, e `prop2val` avrebbe come valore la stringa `val2`.

# --instructions--

Leggi i valori delle proprietà di `testObj` usando la notazione a punto. Imposta la variabile `hatValue` uguale alla proprietà `hat` dell’oggetto e imposta la variabile `shirtValue` uguale alla proprietà `shirt` dell’oggetto.

# --hints--

`hatValue` dovrebbe essere una stringa

```js
assert(typeof hatValue === 'string');
```

Il valore di `hatValue` dovrebbe essere la stringa `ballcap`

```js
assert(hatValue === 'ballcap');
```

`shirtValue` dovrebbe essere una stringa

```js
assert(typeof shirtValue === 'string');
```

Il valore di `shirtValue` dovrebbe essere la stringa `jersey`

```js
assert(shirtValue === 'jersey');
```

Dovresti usare la notazione a punto due volte

```js
assert(__helpers.removeJSComments(code).match(/testObj\.\w+/g).length > 1);
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  hat: "ballcap",
  shirt: "jersey",
  shoes: "cleats"
};

// Only change code below this line
const hatValue = testObj;      // Change this line
const shirtValue = testObj;    // Change this line
```

# --solutions--

```js
const testObj = {
  hat: "ballcap",
  shirt: "jersey",
  shoes: "cleats"
};

const hatValue = testObj.hat;
const shirtValue = testObj.shirt;
```
