---
id: 56533eb9ac21ba0edf2244b7
title: "Concatenare stringhe con l'operatore più"
challengeType: 1
forumTopicId: 16802
dashedName: concatenating-strings-with-plus-operator
---

# --description--

In JavaScript, quando l'operatore `+` viene usato con un valore `String`, si chiama operatore di <dfn>concatenazione</dfn>. Puoi creare una nuova stringa a partire da altre stringhe <dfn>concatenandole</dfn> insieme.

**Esempio**

```js
'My name is Alan,' + ' I concatenate.'
```

**Nota:** Fai attenzione agli spazi. La concatenazione non aggiunge spazi tra le stringhe concatenate, quindi dovrai aggiungerli tu stesso.

Esempio:

```js
const ourStr = "I come first. " + "I come second.";
```

La stringa `I come first. I come second.` verrebbe mostrata nella console.
# --instructions--

Crea `myStr` a partire dalle stringhe `This is the start.` e `This is the end.` usando l'operatore `+`. Assicurati di includere uno spazio tra le due stringhe.

# --hints--

`myStr` dovrebbe avere un singolo carattere spazio tra le due stringhe.

```js
assert(/start\. This/.test(myStr));
```

`myStr` dovrebbe avere come valore la stringa `This is the start. This is the end.`

```js
assert(myStr === 'This is the start. This is the end.');
```

Dovresti usare l'operatore `+` per creare `myStr`.

```js
assert(__helpers.removeJSComments(code).match(/(["']).*\1\s*\+\s*(["']).*\2/g));
```

`myStr` dovrebbe essere creato usando la parola chiave `const`.

```js
assert(/const\s+myStr/.test(__helpers.removeJSComments(code)));
```

Dovresti assegnare il risultato alla variabile `myStr`.

```js
assert(/myStr\s*=/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "This is the start. " + "This is the end.";
```
