---
id: 56533eb9ac21ba0edf2244b8
title: Concatenare stringhe con l'operatore plus equals
challengeType: 1
forumTopicId: 16803
dashedName: concatenating-strings-with-the-plus-equals-operator
---

# --description--

Possiamo anche usare l'operatore `+=` per <dfn>concatenare</dfn> una stringa alla fine di una variabile stringa esistente. Questo può essere molto utile per spezzare una stringa lunga su più righe.

**Nota:** Fai attenzione agli spazi. La concatenazione non aggiunge spazi tra le stringhe concatenate, quindi dovrai aggiungerli tu stesso.

Esempio:

```js
let ourStr = "I come first. ";
ourStr += "I come second.";
```

`ourStr` ora ha come valore la stringa `I come first. I come second.`.

# --instructions--

Costruisci `myStr` su più righe concatenando queste due stringhe: `This is the first sentence.` e `This is the second sentence.` usando l'operatore `+=`. Usa l'operatore `+=` in modo simile a come mostrato nell'esempio e assicurati di includere uno spazio tra le due stringhe. Inizia assegnando la prima stringa a `myStr`, poi aggiungi la seconda stringa.

# --hints--

`myStr` dovrebbe avere un singolo carattere spazio tra le due stringhe.

```js
assert(/sentence\. This/.test(myStr));
```

`myStr` dovrebbe avere come valore la stringa `This is the first sentence. This is the second sentence.`

```js
assert(myStr === 'This is the first sentence. This is the second sentence.');
```

Dovresti usare l'operatore `+=` per costruire `myStr`.

```js
assert(__helpers.removeJSComments(code).match(/myStr\s*\+=\s*(["']).*\1/g));
```

# --seed--

## --seed-contents--

```js
let myStr;
```

# --solutions--

```js
let myStr = "This is the first sentence. ";
myStr += "This is the second sentence.";
```
