---
id: 56533eb9ac21ba0edf2244bd
title: Passare valori alle funzioni con argomenti
challengeType: 1
forumTopicId: 18254
dashedName: passing-values-to-functions-with-arguments
---

# --description--

<dfn>I parametri</dfn> sono variabili che fungono da segnaposto per i valori che devono essere forniti a una funzione quando viene chiamata. Quando si definisce una funzione, di solito la si definisce insieme a uno o più parametri. I valori effettivi che vengono forniti (o <dfn>"passati"</dfn>) a una funzione quando viene chiamata sono noti come <dfn>argomenti</dfn>.

Ecco una funzione con due parametri, `param1` e `param2`:

```js
function testFun(param1, param2) {
  console.log(param1, param2);
}
```

Poi possiamo chiamare `testFun` così: `testFun("Hello", "World");`. Abbiamo passato due argomenti stringa, `Hello` e `World`. All'interno della funzione, `param1` sarà uguale alla stringa `Hello` e `param2` sarà uguale alla stringa `World`. Nota che potresti chiamare di nuovo `testFun` con argomenti diversi e i parametri assumerebbero il valore dei nuovi argomenti.

# --instructions--

<ol><li>Crea una funzione chiamata <code>functionWithArgs</code> che accetti due argomenti e stampi la loro somma nella console di sviluppo.</li><li>Chiama la funzione con due numeri come argomenti.</li></ol>

# --before-each--

```js
var logOutput;
```

# --hints--

`functionWithArgs` dovrebbe essere una funzione.

```js
assert(typeof functionWithArgs === 'function');
```

`functionWithArgs(1,2)` dovrebbe stampare `3`.

```js
console.log = function (message) {
  logOutput = message
};
functionWithArgs(1, 2);
assert(logOutput == 3);
```

`functionWithArgs(7,9)` dovrebbe stampare `16`.

```js
console.log = function (message) {
  logOutput = message
};
functionWithArgs(7, 9);
assert(logOutput == 16);
```

Dovresti chiamare `functionWithArgs` con due numeri dopo averla definita.

```js
assert(
  /functionWithArgs\([-+]?\d*\.?\d*,[-+]?\d*\.?\d*\)/.test(
    __helpers.removeJSComments(code).replace(/\s/g, '')
  )
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function functionWithArgs(a, b) {
  console.log(a + b);
}
functionWithArgs(10, 5);
```
