---
id: 56533eb9ac21ba0edf2244c3
title: Assegnazione con un valore di ritorno
challengeType: 1
forumTopicId: 16658
dashedName: assignment-with-a-returned-value
---

# --description--

Se ricordi dalla nostra discussione su <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/storing-values-with-the-assignment-operator" target="_blank" rel="noopener noreferrer nofollow">Memorizzare valori con l'operatore di assegnazione</a>, tutto ciò che si trova a destra del segno di uguale viene risolto prima che il valore venga assegnato. Questo significa che possiamo prendere il valore di ritorno di una funzione e assegnarlo a una variabile.

Supponiamo di aver definito una funzione `sum` che somma due numeri.

```js
ourSum = sum(5, 12);
```

Chiamare la funzione `sum` con gli argomenti `5` e `12` produce un valore di ritorno pari a `17`. Questo valore di ritorno viene assegnato alla variabile `ourSum`.

# --instructions--

Chiama la funzione `processArg` con un argomento pari a `7` e assegna il suo valore di ritorno alla variabile `processed`.

# --hints--

`processed` dovrebbe avere un valore pari a `2`

```js
assert(processed === 2);
```

Dovresti assegnare `processArg` a `processed`

```js
assert(/processed\s*=\s*processArg\(\s*7\s*\)/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
let processed = 0;

function processArg(num) {
  return (num + 3) / 5;
}

// Only change code below this line

```

# --solutions--

```js
var processed = 0;

function processArg(num) {
  return (num + 3) / 5;
}

processed = processArg(7);
```
