---
id: 56533eb9ac21ba0edf2244ed
title: Aggiungere variabili alle stringhe
challengeType: 1
forumTopicId: 16656
dashedName: appending-variables-to-strings
---

# --description--

Proprio come possiamo costruire una stringa su più righe usando letterali di stringa <dfn>literals</dfn>, possiamo anche aggiungere variabili a una stringa usando l'operatore più uguale (`+=`).

Esempio:

```js
const anAdjective = "awesome!";
let ourStr = "freeCodeCamp is ";
ourStr += anAdjective;
```

`ourStr` avrebbe il valore `freeCodeCamp is awesome!`.

# --instructions--

Imposta `someAdjective` su una stringa di almeno 3 caratteri e aggiungila a `myStr` usando l'operatore `+=`.

# --hints--

`someAdjective` dovrebbe essere impostato su una stringa lunga almeno 3 caratteri.

```js
assert(typeof someAdjective !== 'undefined' && someAdjective.length > 2);
```

Dovresti aggiungere `someAdjective` a `myStr` usando l'operatore `+=`.

```js
assert(__helpers.removeJSComments(code).match(/myStr\s*\+=\s*someAdjective\s*/).length > 0);
```

# --seed--

## --seed-contents--

```js
// Change code below this line
const someAdjective = "";
let myStr = "Learning to code is ";
```

# --solutions--

```js
const someAdjective = "neat";
let myStr = "Learning to code is ";
myStr += someAdjective;
```
