---
id: 56533eb9ac21ba0edf2244c9
title: Consultare le proprietà di un oggetto con variabili
challengeType: 1
forumTopicId: 16165
dashedName: accessing-object-properties-with-variables
---

# --description--

Un altro uso della notazione a parentesi sugli oggetti è consultare una proprietà che è archiviata come valore di una variabile. Questo può essere molto utile per scorrere le proprietà di un oggetto o quando si consulta una tabella di riferimento.

Ecco un esempio di come usare una variabile per consultare una proprietà:

```js
const dogs = {
  Fido: "Mutt",
  Hunter: "Doberman",
  Snoopie: "Beagle"
};

const myDog = "Hunter";
const myBreed = dogs[myDog];
console.log(myBreed);
```

La stringa `Doberman` verrebbe mostrata nella console.

Nota che *non* usiamo le virgolette intorno al nome della variabile quando la usiamo per consultare la proprietà perché stiamo usando il *valore* della variabile, non il *nome*.

# --instructions--

Assegna alla variabile `playerNumber` il valore `16`. Poi, usa la variabile per consultare il nome del giocatore e assegnalo a `player`.

# --hints--

`playerNumber` dovrebbe essere un numero

```js
assert(typeof playerNumber === 'number');
```

La variabile `player` dovrebbe essere una stringa

```js
assert(typeof player === 'string');
```

Il valore di `player` dovrebbe essere la stringa `Montana`

```js
assert(player === 'Montana');
```

Dovresti usare la notazione a parentesi per consultare `testObj`

```js
assert(/testObj\s*?\[.*?\]/.test(__helpers.removeJSComments(code)));
```

Non dovresti assegnare direttamente il valore `Montana` alla variabile `player`

```js
assert(!__helpers.removeJSComments(code).match(/player\s*=\s*"|\'\s*Montana\s*"|\'\s*;/gi));
```

Dovresti usare la variabile `playerNumber` nella tua notazione a parentesi

```js
assert(/testObj\s*?\[\s*playerNumber\s*\]/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  12: "Namath",
  16: "Montana",
  19: "Unitas"
};

// Only change code below this line
const playerNumber = 42;  // Change this line
const player = testObj;   // Change this line
```

# --solutions--

```js
const testObj = {
  12: "Namath",
  16: "Montana",
  19: "Unitas"
};
const playerNumber = 16;
const player = testObj[playerNumber];
```
