---
id: 56533eb9ac21ba0edf2244bc
title: Lista della spesa
challengeType: 1
forumTopicId: 18280
dashedName: shopping-list
---

# --description--

Crea una lista della spesa nella variabile `myList`. La lista deve essere un array multidimensionale contenente diversi sotto-array.

Il primo elemento di ogni sotto-array deve contenere una stringa con il nome dell'elemento. Il secondo elemento deve essere un numero che rappresenta la quantità, ad esempio

```js
["Chocolate Bar", 15]
```

La lista deve contenere almeno 5 sotto-array.

# --hints--

`myList` dovrebbe essere un array.

```js
assert(Array.isArray(myList));
```

I primi elementi di ciascuno dei tuoi sotto-array devono essere tutti stringhe.

```js
assert(myList.every(elem => elem && typeof elem[0] === 'string'));
```

I secondi elementi di ciascuno dei tuoi sotto-array devono essere tutti numeri.

```js
assert(myList.every(elem => elem && typeof elem[1] === 'number'));
```

Dovresti avere almeno 5 elementi nella tua lista.

```js
assert(myList.length > 4);
```

# --seed--

## --seed-contents--

```js
const myList = [];
```

# --solutions--

```js
const myList = [
  ["Candy", 10],
  ["Potatoes", 12],
  ["Eggs", 12],
  ["Catfood", 1],
  ["Toads", 9]
];
```
