---
id: 56533eb9ac21ba0edf2244b9
title: Costruire stringhe con variabili
challengeType: 1
forumTopicId: 16805
dashedName: constructing-strings-with-variables
---

# --description--

A volte dovrai costruire una stringa. Usando l'operatore di concatenazione (`+`), puoi inserire una o più variabili in una stringa che stai costruendo.

Esempio:

```js
const ourName = "freeCodeCamp";
const ourStr = "Hello, our name is " + ourName + ", how are you?";
```

`ourStr` avrebbe un valore della stringa `Hello, our name is freeCodeCamp, how are you?`.

# --instructions--

Imposta `myName` su una stringa uguale al tuo nome e costruisci `myStr` con `myName` tra le stringhe `My name is ` e ` and I am well!`

# --hints--

`myName` dovrebbe essere impostato su una stringa di almeno 3 caratteri.

```js
assert(typeof myName !== 'undefined' && myName.length > 2);
```

Dovresti usare due operatori `+` per costruire `myStr` con `myName` al suo interno.

```js
assert(__helpers.removeJSComments(code).match(/["']\s*\+\s*myName\s*\+\s*["']/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const myName = "";
const myStr = "";
```

# --solutions--

```js
const myName = "Bob";
const myStr = "My name is " + myName + " and I am well!";
```
