---
id: 56533eb9ac21ba0edf2244de
title: "aggiungere un'opzione predefinita nelle istruzioni switch"
challengeType: 1
forumTopicId: 16653
dashedName: adding-a-default-option-in-switch-statements
---

# --description--

In un'istruzione `switch` potresti non riuscire a specificare tutti i valori possibili come istruzioni `case`. Invece, puoi aggiungere l'istruzione `default` che verrà eseguita se non viene trovato nessun caso `case` corrispondente. Pensala come l'ultima istruzione `else` in una catena `if/else`.

Un'istruzione `default` dovrebbe essere l'ultimo caso.

```js
switch (num) {
  case value1:
    statement1;
    break;
  case value2:
    statement2;
    break;
...
  default:
    defaultStatement;
    break;
}
```

# --instructions--

Scrivi un'istruzione switch per impostare `answer` per le seguenti condizioni:  
`a` - `apple`  
`b` - `bird`  
`c` - `cat`  
`default` - `stuff`

# --hints--

`switchOfStuff("a")` dovrebbe restituire la stringa `apple`

```js
assert(switchOfStuff('a') === 'apple');
```

`switchOfStuff("b")` dovrebbe restituire la stringa `bird`

```js
assert(switchOfStuff('b') === 'bird');
```

`switchOfStuff("c")` dovrebbe restituire la stringa `cat`

```js
assert(switchOfStuff('c') === 'cat');
```

`switchOfStuff("d")` dovrebbe restituire la stringa `stuff`

```js
assert(switchOfStuff('d') === 'stuff');
```

`switchOfStuff(4)` dovrebbe restituire la stringa `stuff`

```js
assert(switchOfStuff(4) === 'stuff');
```

Non dovresti usare istruzioni `if` o `else`

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

Dovresti usare un'istruzione `default`

```js
assert(switchOfStuff('string-to-trigger-default-case') === 'stuff');
```

Dovresti avere almeno 3 istruzioni `break`

```js
assert(__helpers.removeJSComments(code).match(/break/g).length > 2);
```

# --seed--

## --seed-contents--

```js
function switchOfStuff(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

switchOfStuff(1);
```

# --solutions--

```js
function switchOfStuff(val) {
  let answer = "";

  switch(val) {
    case "a":
      answer = "apple";
      break;
    case "b":
      answer = "bird";
      break;
    case "c":
      answer = "cat";
      break;
    default:
      answer = "stuff";
  }
  return answer;
}
```
