---
id: 56533eb9ac21ba0edf2244dd
title: Selezionare tra molte opzioni con istruzioni switch
challengeType: 1
forumTopicId: 18277
dashedName: selecting-from-many-options-with-switch-statements
---

# --description--

Se devi confrontare un valore con molte opzioni, puoi usare un’istruzione <dfn>switch</dfn>. Un’istruzione `switch` confronta il valore con le istruzioni <dfn>case</dfn> che definiscono vari valori possibili. Qualsiasi istruzione JavaScript valida può essere eseguita all’interno di un blocco <dfn>case</dfn> e verrà eseguita dal primo valore `case` corrispondente fino a quando non si incontra un `break`.

Ecco un esempio di un’istruzione `switch`:

```js
switch (fruit) {
  case "apple":
    console.log("The fruit is an apple");
    break;
  case "orange":
    console.log("The fruit is an orange");
    break;
}
```

I valori `case` sono testati con uguaglianza stretta (`===`). Il `break` dice a JavaScript di smettere di eseguire le istruzioni. Se il `break` viene omesso, verrà eseguita l’istruzione successiva.

# --instructions--

Scrivi un’istruzione switch che verifichi `val` e imposti `answer` per le seguenti condizioni:  
`1` - `alpha`  
`2` - `beta`  
`3` - `gamma`  
`4` - `delta`

# --hints--

`caseInSwitch(1)` dovrebbe avere il valore della stringa `alpha`

```js
assert(caseInSwitch(1) === 'alpha');
```

`caseInSwitch(2)` dovrebbe avere il valore della stringa `beta`

```js
assert(caseInSwitch(2) === 'beta');
```

`caseInSwitch(3)` dovrebbe avere il valore della stringa `gamma`

```js
assert(caseInSwitch(3) === 'gamma');
```

`caseInSwitch(4)` dovrebbe avere il valore della stringa `delta`

```js
assert(caseInSwitch(4) === 'delta');
```

Non dovresti usare istruzioni `if` o `else`

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

Dovresti avere almeno 3 istruzioni `break`

```js
assert(__helpers.removeJSComments(code).match(/break/g).length > 2);
```

# --seed--

## --seed-contents--

```js
function caseInSwitch(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

caseInSwitch(1);
```

# --solutions--

```js
function caseInSwitch(val) {
  let answer = "";

  switch (val) {
    case 1:
      answer = "alpha";
      break;
    case 2:
      answer = "beta";
      break;
    case 3:
      answer = "gamma";
      break;
    case 4:
      answer = "delta";
  }
  return answer;
}
```
