---
id: 56533eb9ac21ba0edf2244ba
title: Comprendere l'immutabilità delle stringhe
challengeType: 1
forumTopicId: 18331
dashedName: understand-string-immutability
---

# --description--

In JavaScript, i valori `String` sono <dfn>immutabili</dfn>, il che significa che non possono essere modificati una volta creati.

Ad esempio, il seguente codice genererà un errore perché la lettera `B` nella stringa `Bob` non può essere cambiata nella lettera `J`:

```js
let myStr = "Bob";
myStr[0] = "J";
```

Nota che questo *non* significa che `myStr` non possa essere riassegnato. L’unico modo per cambiare `myStr` sarebbe assegnargli un nuovo valore, come in questo esempio:

```js
let myStr = "Bob";
myStr = "Job";
```

# --instructions--

Correggi l’assegnazione a `myStr` in modo che contenga il valore stringa di `Hello World` usando l’approccio mostrato nell’esempio sopra.

# --hints--

`myStr` dovrebbe avere un valore della stringa `Hello World`.

```js
assert(myStr === 'Hello World');
```

Non dovresti modificare il codice sopra il commento specificato.

```js
assert(/myStr = "Jello World"/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
let myStr = "Jello World";

// Only change code below this line
myStr[0] = "H"; // Change this line
// Only change code above this line
```

# --solutions--

```js
let myStr = "Jello World";
myStr = "Hello World";
```
