---
id: 56533eb9ac21ba0edf2244ae
title: Trovare il resto in JavaScript
challengeType: 1
forumTopicId: 18184
dashedName: finding-a-remainder-in-javascript
---

# --description--

L’<dfn>operatore resto</dfn> `%` restituisce il resto della divisione tra due numeri.

**Esempio**

<pre>
5 % 2 = 1
5 / 2 = 2 resto 1
2 * 2 = 4
5 - 4 = 1
</pre>

**Utilizzo**  
In matematica, un numero può essere verificato come pari o dispari controllando il resto della divisione del numero per `2`. I numeri pari hanno un resto di `0`, mentre i numeri dispari un resto di `1`.

<pre>
17 % 2 = 1
48 % 2 = 0
</pre>

**Nota:** L’<dfn>operatore resto</dfn> è talvolta chiamato erroneamente operatore modulo. È molto simile al modulo, ma non funziona correttamente con numeri negativi.

# --instructions--

Imposta `remainder` uguale al resto di `11` diviso per `3` usando l’<dfn>operatore resto</dfn> (`%`).

# --hints--

La variabile `remainder` dovrebbe essere inizializzata

```js
assert(/(const|let|var)\s+?remainder/.test(__helpers.removeJSComments(code)));
```

Il valore di `remainder` dovrebbe essere `2`

```js
assert(remainder === 2);
```

Dovresti usare l’operatore `%`

```js
assert(/\s+?remainder\s*?=\s*?.*%.*;?/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const remainder = 0;
```

# --solutions--

```js
const remainder = 11 % 3;
```
