---
id: 56533eb9ac21ba0edf2244cc
title: Consultare oggetti annidati
challengeType: 1
forumTopicId: 16161
dashedName: accessing-nested-objects
---

# --description--

Le sotto-proprietà degli oggetti possono essere consultate concatenando la notazione a punto o la notazione a parentesi.

Ecco un oggetto annidato:

```js
const ourStorage = {
  "desk": {
    "drawer": "stapler"
  },
  "cabinet": {
    "top drawer": { 
      "folder1": "a file",
      "folder2": "secrets"
    },
    "bottom drawer": "soda"
  }
};

ourStorage.cabinet["top drawer"].folder2;
ourStorage.desk.drawer;
```

`ourStorage.cabinet["top drawer"].folder2` sarebbe la stringa `secrets`, e `ourStorage.desk.drawer` sarebbe la stringa `stapler`.

# --instructions--

Consulta l’oggetto `myStorage` e assegna il contenuto della proprietà `glove box` alla variabile `gloveBoxContents`. Usa la notazione a punto per tutte le proprietà dove possibile, altrimenti usa la notazione a parentesi.

# --hints--

`gloveBoxContents` dovrebbe essere uguale alla stringa `maps`.

```js
assert(gloveBoxContents === 'maps');
```

Il tuo codice dovrebbe usare la notazione a punto, dove possibile, per consultare `myStorage`.

```js
assert.match(code, /myStorage\.car\.inside/);
```

`gloveBoxContents` dovrebbe essere comunque dichiarato con `const`.

```js
assert.match(code, /const\s+gloveBoxContents\s*=/);
```

Non dovresti modificare l’oggetto `myStorage`.

```js
const expectedMyStorage = {
  "car":{
    "inside":{
      "glove box":"maps",
      "passenger seat":"crumbs"
    },
    "outside":{
      "trunk":"jack"
    }
  }
};
assert.deepStrictEqual(myStorage, expectedMyStorage);
```

# --seed--

## --seed-contents--

```js
const myStorage = {
  "car": {
    "inside": {
      "glove box": "maps",
      "passenger seat": "crumbs"
     },
    "outside": {
      "trunk": "jack"
    }
  }
};

const gloveBoxContents = undefined;
```

# --solutions--

```js
const myStorage = {
  "car":{
    "inside":{
      "glove box":"maps",
      "passenger seat":"crumbs"
    },
    "outside":{
      "trunk":"jack"
    }
  }
};
const gloveBoxContents = myStorage.car.inside["glove box"];
```
