---
id: 56533eb9ac21ba0edf2244be
title: Ambito globale e funzioni
challengeType: 1
forumTopicId: 18193
dashedName: global-scope-and-functions
---

# --description--

In JavaScript, <dfn>ambito</dfn> indica la visibilità delle variabili. Le variabili definite all'esterno di un blocco di funzione hanno ambito <dfn>globale</dfn>. Questo significa che possono essere viste ovunque nel tuo codice JavaScript.

Le variabili dichiarate senza le parole chiave `let` o `const` vengono create automaticamente nell'ambito `global`. Questo può causare conseguenze indesiderate in altre parti del codice o quando si esegue nuovamente una funzione. Dovresti sempre dichiarare le tue variabili con `let` o `const`.

# --instructions--

Usando `let` o `const`, dichiara una variabile globale chiamata `myGlobal` all'esterno di qualsiasi funzione. Inizializzala con un valore di `10`.

All'interno della funzione `fun1`, assegna `5` a `oopsGlobal` ***senza*** usare le parole chiave `var`, `let` o `const`.

# --before-each--

```js
var oopsGlobal;
```

# --hints--

`myGlobal` dovrebbe essere definito

```js
assert(typeof myGlobal != 'undefined');
```

`myGlobal` dovrebbe avere un valore di `10`

```js
assert(myGlobal === 10);
```

`myGlobal` dovrebbe essere dichiarato usando le parole chiave `let` o `const`

```js
assert(/(let|const)\s+myGlobal/.test(__helpers.removeJSComments(code)));
```

`oopsGlobal` dovrebbe essere una variabile globale e avere un valore di `5`

```js
fun1();
assert(typeof oopsGlobal != 'undefined');
```

# --seed--

## --seed-contents--

```js
// Declare the myGlobal variable below this line


function fun1() {
  // Assign 5 to oopsGlobal here

}

// Only change code above this line

function fun2() {
  let output = "";
  if (typeof myGlobal != "undefined") {
    output += "myGlobal: " + myGlobal;
  }
  if (typeof oopsGlobal != "undefined") {
    output += " oopsGlobal: " + oopsGlobal;
  }
  console.log(output);
}
```

# --solutions--

```js
const myGlobal = 10;

function fun1() {
  oopsGlobal = 5;
}

function fun2() {
  let output = "";
  if(typeof myGlobal != "undefined") {
    output += "myGlobal: " + myGlobal;
  }
  if(typeof oopsGlobal != "undefined") {
    output += " oopsGlobal: " + oopsGlobal;
  }
  console.log(output);
}
```
