---
id: 56533eb9ac21ba0edf2244bb
title: Frasi con spazi vuoti
challengeType: 1
forumTopicId: 18377
dashedName: word-blanks
---

# --description--

Ti vengono fornite frasi con alcune parole mancanti, come nomi, verbi, aggettivi e avverbi. Devi quindi completare gli spazi vuoti con parole a tua scelta in modo che la frase completa abbia senso.

Considera questa frase:

```md
It was really ____, and we ____ ourselves ____.
```

Questa frase ha tre parti mancanti: un aggettivo, un verbo e un avverbio, e possiamo aggiungere parole a nostra scelta per completarla. Possiamo poi assegnare la frase completa a una variabile come segue:

```js
const sentence = "It was really " + "hot" + ", and we " + "laughed" + " ourselves " + "silly" + ".";
```

# --instructions--

In questa sfida, ti forniamo un nome, un verbo, un aggettivo e un avverbio. Devi formare una frase completa usando parole a tua scelta, insieme alle parole che ti forniamo.

Dovrai usare l'operatore di concatenazione `+` per costruire una nuova stringa, usando le variabili fornite: `myNoun`, `myAdjective`, `myVerb` e `myAdverb`. Assegnerai poi la stringa formata alla variabile `wordBlanks`. Non devi modificare le parole assegnate alle variabili.

Dovrai anche considerare gli spazi nella tua stringa, in modo che la frase finale abbia spazi tra tutte le parole. Il risultato dovrebbe essere una frase completa.

# --before-each--

```js
const removeAssignments = str => str
  .replace(/myNoun\s*=\s*["']dog["']/g, '')
  .replace(/myAdjective\s*=\s*["']big["']/g, '')
  .replace(/myVerb\s*=\s*["']ran["']/g, '')
  .replace(/myAdverb\s*=\s*["']quickly["']/g, '');
```

# --hints--

`wordBlanks` dovrebbe essere una stringa.

```js
assert(typeof wordBlanks === 'string');
```

Non devi modificare i valori assegnati a `myNoun`, `myVerb`, `myAdjective` o `myAdverb`.

```js
assert(
  myNoun === 'dog' &&
    myVerb === 'ran' &&
    myAdjective === 'big' &&
    myAdverb === 'quickly'
);
```

Non devi usare direttamente i valori `dog`, `ran`, `big` o `quickly` per creare `wordBlanks`.

```js
const newCode = removeAssignments(__helpers.removeJSComments(code));
assert(
  !/dog/.test(newCode) &&
    !/ran/.test(newCode) &&
    !/big/.test(newCode) &&
    !/quickly/.test(newCode)
);
```

`wordBlanks` dovrebbe contenere tutte le parole assegnate alle variabili `myNoun`, `myVerb`, `myAdjective` e `myAdverb` separate da caratteri non alfanumerici (e qualsiasi altra parola a tua scelta).

```js
assert(
  /\bdog\b/.test(wordBlanks) &&
    /\bbig\b/.test(wordBlanks) &&
    /\bran\b/.test(wordBlanks) &&
    /\bquickly\b/.test(wordBlanks)
);
```

# --seed--

## --seed-contents--

```js
const myNoun = "dog";
const myAdjective = "big";
const myVerb = "ran";
const myAdverb = "quickly";

// Only change code below this line
const wordBlanks = ""; // Change this line
// Only change code above this line
```

# --solutions--

```js
const myNoun = "dog";
const myAdjective = "big";
const myVerb = "ran";
const myAdverb = "quickly";

let wordBlanks = "Once there was a " + myNoun + " which was very " + myAdjective + ". ";
wordBlanks += "It " + myVerb + " " + myAdverb + " around the yard.";
```
