---
id: 56533eb9ac21ba0edf2244a8
title: Memorizzare valori con l'operatore di assegnazione
challengeType: 1
forumTopicId: 18310
dashedName: storing-values-with-the-assignment-operator
---

# --description--

In JavaScript, you can store a value in a variable with the <dfn>assignment</dfn> operator (`=`).

```js
myVariable = 5;
```

Questo assegna il valore `Number` `5` a `myVariable`.

Se sono presenti calcoli a destra dell'operatore `=`, vengono eseguiti prima che il valore venga assegnato alla variabile a sinistra dell'operatore.

```js
var myVar;
myVar = 5;
```

Innanzitutto, questo codice crea una variabile denominata `myVar`. Quindi, il codice assegna `5` a `myVar`. Ora, se `myVar` appare di nuovo nel codice, il programma lo tratterà come se fosse `5`.

# --instructions--

Assegna il valore `7` alla variabile `a`.

# --hints--

Non modificare il codice sopra il commento specificato.

```js
assert(/var a;/.test(__helpers.removeJSComments(code)));
```

`a` dovrebbe avere un valore di 7.

```js
assert(typeof a === 'number' && a === 7);
```

# --seed--

## --seed-contents--

```js
// Setup
var a;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
```
