---
id: 56533eb9ac21ba0edf2244d0
title: "confronto con l'operatore di uguaglianza"
challengeType: 1
forumTopicId: 16784
dashedName: comparison-with-the-equality-operator
---

# --description--

In JavaScript ci sono molti <dfn>operatori di confronto</dfn>. Tutti questi operatori restituiscono un valore booleano `true` o `false`.

L'operatore più semplice è l'operatore di uguaglianza `==`. L'operatore di uguaglianza confronta due valori e restituisce `true` se sono equivalenti o `false` se non lo sono. Nota che uguaglianza è diverso da assegnazione (`=`), che assegna il valore a destra dell'operatore a una variabile a sinistra.

```js
function equalityTest(myVal) {
  if (myVal == 10) {
    return "Equal";
  }
  return "Not Equal";
}
```

Se `myVal` è uguale a `10`, l'operatore di uguaglianza restituisce `true`, quindi il codice nelle parentesi graffe verrà eseguito e la funzione restituirà `Equal`. Altrimenti, la funzione restituirà `Not Equal`. Perché JavaScript possa confrontare due <dfn>tipi di dati</dfn> diversi (per esempio, `numbers` e `strings`), deve convertire un tipo nell'altro. Questo è noto come conversione implicita. Una volta fatto, però, può confrontare i termini come segue:

```js
1   ==  1  // true
1   ==  2  // false
1   == '1' // true
"3" ==  3  // true
```

# --instructions--

Aggiungi l'operatore di uguaglianza nella riga indicata in modo che la funzione restituisca la stringa `Equal` quando `val` è equivalente a `12`.

# --hints--

`testEqual(10)` dovrebbe restituire la stringa `Not Equal`

```js
assert(testEqual(10) === 'Not Equal');
```

`testEqual(12)` dovrebbe restituire la stringa `Equal`

```js
assert(testEqual(12) === 'Equal');
```

`testEqual("12")` dovrebbe restituire la stringa `Equal`

```js
assert(testEqual('12') === 'Equal');
```

Dovresti usare l'operatore `==`

```js
assert(__helpers.removeJSComments(code).match(/==/g) && !__helpers.removeJSComments(code).match(/===/g));
```

# --seed--

## --seed-contents--

```js
// Setup
function testEqual(val) {
  if (val) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

testEqual(10);
```

# --solutions--

```js
function testEqual(val) {
  if (val == 12) {
    return "Equal";
  }
  return "Not Equal";
}
```
