---
id: 56533eb9ac21ba0edf2244cd
title: Consultare array annidati
challengeType: 1
forumTopicId: 16160
dashedName: accessing-nested-arrays
---

# --description--

Come abbiamo visto in esempi precedenti, gli oggetti possono contenere sia oggetti annidati sia array annidati. Similmente a come si consultano gli oggetti annidati, la notazione a parentesi per gli array può essere concatenata per consultare array annidati.

Ecco un esempio di come consultare un array annidato:

```js
const ourPets = [
  {
    animalType: "cat",
    names: [
      "Meowzer",
      "Fluffy",
      "Kit-Cat"
    ]
  },
  {
    animalType: "dog",
    names: [
      "Spot",
      "Bowser",
      "Frankie"
    ]
  }
];

ourPets[0].names[1];
ourPets[1].names[0];
```

`ourPets[0].names[1]` sarebbe la stringa `Fluffy`, e `ourPets[1].names[0]` sarebbe la stringa `Spot`.

# --instructions--

Usando la notazione a punto e a parentesi, imposta la variabile `secondTree` al secondo elemento nell'array `list` del secondo oggetto nell'array `myPlants`.

# --hints--

`secondTree` dovrebbe essere uguale alla stringa `pine`.

```js
assert(secondTree === 'pine');
```

Il tuo codice dovrebbe usare la notazione a punto e a parentesi per consultare `myPlants`.

```js
assert(/=\s*myPlants\[1\].list\[1\]/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myPlants = [
  {
    type: "flowers",
    list: [
      "rose",
      "tulip",
      "dandelion"
    ]
  },
  {
    type: "trees",
    list: [
      "fir",
      "pine",
      "birch"
    ]
  }
];

const secondTree = "";
```

# --solutions--

```js
const myPlants = [
  {
    type: "flowers",
    list: [
      "rose",
      "tulip",
      "dandelion"
    ]
  },
  {
    type: "trees",
    list: [
      "fir",
      "pine",
      "birch"
    ]
  }
];

const secondTree = myPlants[1].list[1];
```
