---
id: 56533eb9ac21ba0edf2244cb
title: Manipolare oggetti complessi
challengeType: 1
forumTopicId: 18208
dashedName: manipulating-complex-objects
---

# --description--

A volte potresti voler archiviare dati in una <dfn>struttura dati</dfn> flessibile. Un oggetto JavaScript è un modo per gestire dati flessibili. Permettono combinazioni arbitrarie di <dfn>stringhe</dfn>, <dfn>numeri</dfn>, <dfn>booleani</dfn>, <dfn>array</dfn>, <dfn>funzioni</dfn> e <dfn>oggetti</dfn>.

Ecco un esempio di una struttura dati complessa:

```js
const ourMusic = [
  {
    "artist": "Daft Punk",
    "title": "Homework",
    "release_year": 1997,
    "formats": [ 
      "CD", 
      "Cassette", 
      "LP"
    ],
    "gold": true
  }
];
```

Questo è un array che contiene un oggetto al suo interno. L’oggetto ha vari pezzi di <dfn>metadati</dfn> su un album. Ha anche un array `formats` annidato. Se vuoi aggiungere altri record di album, puoi farlo aggiungendo record all’array di primo livello. Gli oggetti contengono dati in una proprietà, che ha un formato chiave-valore. Nell’esempio sopra, `"artist": "Daft Punk"` è una proprietà che ha come chiave `artist` e come valore `Daft Punk`.

**Nota:** Devi mettere una virgola dopo ogni oggetto nell’array, a meno che non sia l’ultimo oggetto nell’array.

# --instructions--

Aggiungi un nuovo album all’array `myMusic`. Aggiungi le stringhe `artist` e `title`, il numero `release_year` e un array `formats` di stringhe.

# --hints--

`myMusic` dovrebbe essere un array

```js
assert(Array.isArray(myMusic));
```

`myMusic` dovrebbe avere almeno due elementi

```js
assert(myMusic.length > 1);
```

Gli elementi nell’array `myMusic` dovrebbero essere oggetti

```js
myMusic.forEach(object => {assert.typeOf(object, 'object')})
```

Il tuo oggetto in `myMusic` dovrebbe avere almeno 4 proprietà

```js
myMusic.forEach(object => {assert(Object.keys(object).length > 3); });
```

Il tuo oggetto in `myMusic` dovrebbe contenere la proprietà `artist` che è una stringa

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['artist']);
  assert.typeOf(object.artist, 'string')
})
```

Il tuo oggetto in `myMusic` dovrebbe contenere la proprietà `title` che è una stringa

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['title']);
  assert.typeOf(object.title, 'string')
})
```

Il tuo oggetto in `myMusic` dovrebbe contenere la proprietà `release_year` che è un numero

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['release_year']);
  assert.typeOf(object.release_year, 'number')
})
```

Il tuo oggetto in `myMusic` dovrebbe contenere una proprietà `formats` che è un array

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['formats']);
  assert.typeOf(object.formats, 'array')
})
```

`formats` dovrebbe essere un array di stringhe con almeno due elementi

```js
myMusic.forEach(object => {
  object.formats.forEach(format => {
    assert.typeOf(format, 'string')
  });
  assert.isAtLeast(object.formats.length, 2)
})
```

# --seed--

## --seed-contents--

```js
const myMusic = [
  {
    "artist": "Billy Joel",
    "title": "Piano Man",
    "release_year": 1973,
    "formats": [
      "CD",
      "8T",
      "LP"
    ],
    "gold": true
  }
];
```

# --solutions--

```js
const myMusic = [
  {
    "artist": "Billy Joel",
    "title": "Piano Man",
    "release_year": 1973,
    "formats": [
      "CS",
      "8T",
      "LP" ],
    "gold": true
  },
  {
    "artist": "ABBA",
    "title": "Ring Ring",
    "release_year": 1973,
    "formats": [
      "CS",
      "8T",
      "LP",
    "CD",
  ]
  }
];
```
