---
id: 56533eb9ac21ba0edf2244b9
title: Costruire stringhe con le variabili
challengeType: 1
forumTopicId: 16805
dashedName: constructing-strings-with-variables
---

# --description--

Sometimes you will need to build a string. By using the concatenation operator (`+`), you can insert one or more variables into a string you're building.

Esempio:

```js
const ourName = "freeCodeCamp";
const ourStr = "Hello, our name is " + ourName + ", how are you?";
```

`ourStr` avrà un valore stringa `Hello, our name is freeCodeCamp, how are you?`.

# --instructions--

Assegna una stringa con valore pari al tuo nome a `myName` e crea `myStr` con `myName` tra le stringhe `My name is` e `and I am well!`

# --hints--

`myName` dovrebbe avere il valore di una stringa lunga almeno 3 caratteri.

```js
assert(typeof myName !== 'undefined' && myName.length > 2);
```

Dovresti usare due operatori `+` per costruire `myStr` con `myName` al suo interno.

```js
assert(__helpers.removeJSComments(code).match(/["']\s*\+\s*myName\s*\+\s*["']/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const myName = "";
const myStr = "";
```

# --solutions--

```js
const myName = "Bob";
const myStr = "My name is " + myName + " and I am well!";
```
