---
id: 56533eb9ac21ba0edf2244e1
title: Annidare i loop for
challengeType: 1
forumTopicId: 18248
dashedName: nesting-for-loops
---

# --description--

Se hai un array multidimensionale, puoi usare la stessa logica del waypoint precedente per iterare sia l'array principale sia i sotto-array. Ecco un esempio:

```js
const arr = [
  [1, 2], [3, 4], [5, 6]
];

for (let i = 0; i < arr.length; i++) {
  for (let j = 0; j < arr[i].length; j++) {
    console.log(arr[i][j]);
  }
}
```

Questo stampa ogni sotto-elemento in `arr` uno alla volta. Nota che per il loop interno controlliamo la `.length` di `arr[i]`, dato che `arr[i]` è a sua volta un array.

# --instructions--

Modifica la funzione `multiplyAll` in modo che restituisca il prodotto di tutti i numeri presenti nei sotto-array di `arr`.

# --hints--

`multiplyAll([[1], [2], [3]])` dovrebbe restituire `6`

```js
assert(multiplyAll([[1], [2], [3]]) === 6);
```

`multiplyAll([[1, 2], [3, 4], [5, 6, 7]])` dovrebbe restituire `5040`

```js
assert(
  multiplyAll([
    [1, 2],
    [3, 4],
    [5, 6, 7]
  ]) === 5040
);
```

`multiplyAll([[5, 1], [0.2, 4, 0.5], [3, 9]])` dovrebbe restituire `54`

```js
assert(
  multiplyAll([
    [5, 1],
    [0.2, 4, 0.5],
    [3, 9]
  ]) === 54
);
```

# --seed--

## --seed-contents--

```js
function multiplyAll(arr) {
  let product = 1;
  // Only change code below this line

  // Only change code above this line
  return product;
}

multiplyAll([[1, 2], [3, 4], [5, 6, 7]]);
```

# --solutions--

```js
function multiplyAll(arr) {
  let product = 1;
  for (let i = 0; i < arr.length; i++) {
    for (let j = 0; j < arr[i].length; j++) {
      product *= arr[i][j];
    }
  }
  return product;
}
```
