---
id: 56533eb9ac21ba0edf2244c8
title: Consultare le proprietà di un oggetto con la notazione a parentesi
challengeType: 1
forumTopicId: 16163
dashedName: accessing-object-properties-with-bracket-notation
---

# --description--

Il secondo modo per consultare le proprietà di un oggetto è la notazione a parentesi (`[]`). Se la proprietà dell'oggetto che vuoi consultare contiene uno spazio nel nome, dovrai usare la notazione a parentesi.

Tuttavia, puoi comunque usare la notazione a parentesi anche per proprietà di oggetti senza spazi.

Ecco un esempio di come usare la notazione a parentesi per leggere una proprietà di un oggetto:

```js
const myObj = {
  "Space Name": "Kirk",
  "More Space": "Spock",
  "NoSpace": "USS Enterprise"
};

myObj["Space Name"];
myObj['More Space'];
myObj["NoSpace"];
```

`myObj["Space Name"]` sarebbe la stringa `Kirk`, `myObj['More Space']` sarebbe la stringa `Spock`, e `myObj["NoSpace"]` sarebbe la stringa `USS Enterprise`.

Nota che i nomi delle proprietà con spazi devono essere racchiusi tra virgolette (singole o doppie).

# --instructions--

Leggi i valori delle proprietà `an entree` e `the drink` di `testObj` usando la notazione a parentesi e assegnali rispettivamente a `entreeValue` e `drinkValue`.

# --hints--

`entreeValue` dovrebbe essere una stringa

```js
assert(typeof entreeValue === 'string');
```

Il valore di `entreeValue` dovrebbe essere la stringa `hamburger`

```js
assert(entreeValue === 'hamburger');
```

`drinkValue` dovrebbe essere una stringa

```js
assert(typeof drinkValue === 'string');
```

Il valore di `drinkValue` dovrebbe essere la stringa `water`

```js
assert(drinkValue === 'water');
```

Dovresti usare la notazione a parentesi due volte

```js
assert(__helpers.removeJSComments(code).match(/testObj\s*?\[('|")[^'"]+\1\]/g).length > 1);
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  "an entree": "hamburger",
  "my side": "veggies",
  "the drink": "water"
};

// Only change code below this line
const entreeValue = testObj;   // Change this line
const drinkValue = testObj;    // Change this line
```

# --solutions--

```js
const testObj = {
  "an entree": "hamburger",
  "my side": "veggies",
  "the drink": "water"
};
const entreeValue = testObj["an entree"];
const drinkValue = testObj['the drink'];
```
