---
id: 56533eb9ac21ba0edf2244ab
title: Comprendere la distinzione tra maiuscole e minuscole nelle variabili
challengeType: 1
forumTopicId: 18334
dashedName: understanding-case-sensitivity-in-variables
---

# --description--

In JavaScript tutte le variabili e i nomi delle funzioni sono sensibili alle maiuscole e minuscole. Questo significa che la differenza tra maiuscole e minuscole conta.

`MYVAR` non è lo stesso di `MyVar` né di `myvar`. È possibile avere più variabili distinte con lo stesso nome ma con differente uso di maiuscole e minuscole. Si consiglia fortemente, per chiarezza, di *non* usare questa funzionalità del linguaggio.

**Buona pratica**

Scrivi i nomi delle variabili in JavaScript in <dfn>camelCase</dfn>. Nel <dfn>camelCase</dfn>, i nomi di variabili composti da più parole hanno la prima parola in minuscolo e la prima lettera di ogni parola successiva in maiuscolo.

**Esempi:**

```js
var someVariable;
var anotherVariableName;
var thisVariableNameIsSoLong;
```

# --instructions--

Modifica le dichiarazioni e le assegnazioni esistenti in modo che i loro nomi usino il <dfn>camelCase</dfn>.

Non creare nuove variabili.

# --hints--

`studlyCapVar` dovrebbe essere definito e avere il valore `10`.

```js
assert(typeof studlyCapVar !== 'undefined' && studlyCapVar === 10);
```

`properCamelCase` dovrebbe essere definito e avere il valore della stringa `A String`.

```js
assert(
  typeof properCamelCase !== 'undefined' && properCamelCase === 'A String'
);
```

`titleCaseOver` dovrebbe essere definito e avere il valore `9000`.

```js
assert(typeof titleCaseOver !== 'undefined' && titleCaseOver === 9000);
```

`studlyCapVar` dovrebbe usare il camelCase sia nella dichiarazione sia nell’assegnazione.

```js
assert(__helpers.removeJSComments(code).match(/studlyCapVar/g).length === 2);
```

`properCamelCase` dovrebbe usare il camelCase sia nella dichiarazione sia nell’assegnazione.

```js
assert(__helpers.removeJSComments(code).match(/properCamelCase/g).length === 2);
```

`titleCaseOver` dovrebbe usare il camelCase sia nella dichiarazione sia nell’assegnazione.

```js
assert(__helpers.removeJSComments(code).match(/titleCaseOver/g).length === 2);
```

# --seed--

## --seed-contents--

```js
// Variable declarations
var StUdLyCapVaR;
var properCamelCase;
var TitleCaseOver;

// Variable assignments
STUDLYCAPVAR = 10;
PRoperCAmelCAse = "A String";
tITLEcASEoVER = 9000;
```

# --solutions--

```js
var studlyCapVar;
var properCamelCase;
var titleCaseOver;

studlyCapVar = 10;
properCamelCase = "A String";
titleCaseOver = 9000;
```
