---
id: 56592a60ddddeae28f7aa8e1
title: Consultare array multidimensionali con gli indici
challengeType: 1
forumTopicId: 16159
dashedName: access-multi-dimensional-arrays-with-indexes
---

# --description--

Un modo per pensare a un <dfn>array multidimensionale</dfn> è come a un *array di array*. Quando usi le parentesi per consultare il tuo array, il primo gruppo di parentesi si riferisce agli elementi dell'array più esterno (il primo livello), e ogni coppia aggiuntiva di parentesi si riferisce al livello successivo di elementi all'interno.

**Esempio**

```js
const arr = [
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9],
  [[10, 11, 12], 13, 14]
];

const subarray = arr[3];
const nestedSubarray = arr[3][0];
const element = arr[3][0][1];
```

In questo esempio, `subarray` ha il valore `[[10, 11, 12], 13, 14]`,
`nestedSubarray` ha il valore `[10, 11, 12]`, e `element` ha il valore `11`.

**Nota:** Non dovrebbero esserci spazi tra il nome dell'array e le parentesi quadre, come in `array [0][0]` e neanche questo `array [0] [0]` è permesso. Anche se JavaScript riesce a elaborarlo correttamente, questo potrebbe confondere altri programmatori che leggono il tuo codice.

# --instructions--

Usando la notazione a parentesi seleziona un elemento da `myArray` in modo che `myData` sia uguale a `8`.

# --hints--

`myData` dovrebbe essere uguale a `8`.

```js
assert(myData === 8);
```

Dovresti usare la notazione a parentesi per leggere il valore corretto da `myArray`.

```js
assert(/myData=myArray\[2\]\[1\]/.test(__helpers.removeWhiteSpace(__helpers.removeJSComments(code))));
```

# --seed--

## --seed-contents--

```js
const myArray = [
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9],
  [[10, 11, 12], 13, 14],
];

const myData = myArray[0][0];
```

# --solutions--

```js
const myArray = [[1, 2, 3], [4, 5, 6], [7, 8, 9], [[10, 11, 12], 13, 14]];
const myData = myArray[2][1];
```
