---
id: 5675e877dbd60be8ad28edc6
title: Iterare su un array con un loop for
challengeType: 1
forumTopicId: 18216
dashedName: iterate-through-an-array-with-a-for-loop
---

# --description--

Un'attività comune in JavaScript è iterare sul contenuto di un array. Un modo per farlo è con un loop `for`. Questo codice stamperà ogni elemento dell'array `arr` nella console:

```js
const arr = [10, 9, 8, 7, 6];

for (let i = 0; i < arr.length; i++) {
   console.log(arr[i]);
}
```

Ricorda che gli array hanno un indice a base zero, il che significa che l'ultimo indice dell'array è `length - 1`. La nostra condizione per questo loop è `i < arr.length`, che ferma il loop quando `i` è uguale a `length`. In questo caso l'ultima iterazione è `i === 4`, cioè quando `i` diventa uguale a `arr.length - 1` e stampa `6` nella console. Poi `i` aumenta a `5`, e il loop termina perché `i < arr.length` è `false`.

# --instructions--

Dichiara e inizializza una variabile `total` a `0`. Usa un loop `for` per aggiungere il valore di ogni elemento dell'array `myArr` a `total`.

# --hints--

`total` dovrebbe essere dichiarata e inizializzata a 0.

```js
assert(__helpers.removeJSComments(code).match(/(var|let|const)\s*?total\s*=\s*0.*?;?/));
```

`total` dovrebbe essere uguale a 20.

```js
assert(total === 20);
```

Dovresti usare un loop `for` per iterare su `myArr`.

```js
assert(/for\s*\(/g.test(__helpers.removeJSComments(code)) && /myArr\s*\[/g.test(__helpers.removeJSComments(code)));
```

Non dovresti tentare di assegnare direttamente il valore 20 a `total`.

```js
assert(!__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/total[=+-]0*[1-9]+/gm));
```

# --seed--

## --seed-contents--

```js
// Setup
const myArr = [2, 3, 4, 5, 6];

// Only change code below this line

```

# --solutions--

```js
const myArr = [2, 3, 4, 5, 6];
let total = 0;

for (let i = 0; i < myArr.length; i++) {
  total += myArr[i];
}
```
