---
id: 5690307fddb111c6084545d7
title: Ordine logico nelle istruzioni if else
challengeType: 1
forumTopicId: 18228
dashedName: logical-order-in-if-else-statements
---

# --description--

L'ordine è importante nelle istruzioni `if`, `else if`.

La funzione viene eseguita dall'alto verso il basso, quindi devi fare attenzione a quale istruzione viene prima.

Prendi come esempio queste due funzioni.

Ecco la prima:

```js
function foo(x) {
  if (x < 1) {
    return "Less than one";
  } else if (x < 2) {
    return "Less than two";
  } else {
    return "Greater than or equal to two";
  }
}
```

E la seconda inverte semplicemente l'ordine delle istruzioni:

```js
function bar(x) {
  if (x < 2) {
    return "Less than two";
  } else if (x < 1) {
    return "Less than one";
  } else {
    return "Greater than or equal to two";
  }
}
```

Anche se queste due funzioni sembrano quasi identiche, se passiamo un numero a entrambe otteniamo output diversi.

```js
foo(0)
bar(0)
```

`foo(0)` restituirà la stringa `Less than one`, mentre `bar(0)` restituirà la stringa `Less than two`.

# --instructions--

Cambia l'ordine della logica nella funzione in modo che restituisca le istruzioni corrette in tutti i casi.

# --hints--

`orderMyLogic(4)` dovrebbe restituire la stringa `Less than 5`

```js
assert(orderMyLogic(4) === 'Less than 5');
```

`orderMyLogic(6)` dovrebbe restituire la stringa `Less than 10`

```js
assert(orderMyLogic(6) === 'Less than 10');
```

`orderMyLogic(11)` dovrebbe restituire la stringa `Greater than or equal to 10`

```js
assert(orderMyLogic(11) === 'Greater than or equal to 10');
```

# --seed--

## --seed-contents--

```js
function orderMyLogic(val) {
  if (val < 10) {
    return "Less than 10";
  } else if (val < 5) {
    return "Less than 5";
  } else {
    return "Greater than or equal to 10";
  }
}

orderMyLogic(7);
```

# --solutions--

```js
function orderMyLogic(val) {
  if(val < 5) {
    return "Less than 5";
  } else if (val < 10) {
    return "Less than 10";
  } else {
    return "Greater than or equal to 10";
  }
}
```
