---
id: 56bbb991ad1ed5201cd392ce
title: Manipolare gli array con il metodo unshift
challengeType: 1
forumTopicId: 18239
dashedName: manipulate-arrays-with-unshift
---

# --description--

Non solo puoi `shift` elementi dall'inizio di un array, ma puoi anche `unshift` elementi all'inizio di un array, cioè aggiungere elementi davanti all'array.

`.unshift()` funziona esattamente come `.push()`, ma invece di aggiungere l'elemento alla fine dell'array, `unshift()` aggiunge l'elemento all'inizio dell'array.

Esempio:

```js
const ourArray = ["Stimpson", "J", "cat"];
ourArray.shift();
ourArray.unshift("Happy");
```

Dopo il `shift`, `ourArray` avrebbe il valore `["J", "cat"]`. Dopo il `unshift`, `ourArray` avrebbe il valore `["Happy", "J", "cat"]`.

# --instructions--

Aggiungi `["Paul", 35]` all'inizio della variabile `myArray` usando `unshift()`.

# --hints--

`myArray` dovrebbe ora avere `[["Paul", 35], ["dog", 3]]`.

```js
assert(
  (function (d) {
    if (
      typeof d[0] === 'object' &&
      d[0][0] == 'Paul' &&
      d[0][1] === 35 &&
      d[1][0] != undefined &&
      d[1][0] == 'dog' &&
      d[1][1] != undefined &&
      d[1][1] == 3
    ) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["dog", 3]];
myArray.shift();

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["dog", 3]];
myArray.shift();
myArray.unshift(["Paul", 35]);
```
