---
id: 56bbb991ad1ed5201cd392cd
title: Manipolare gli array con il metodo shift
challengeType: 1
forumTopicId: 18238
dashedName: manipulate-arrays-with-shift
---

# --description--

`pop()` rimuove sempre l'ultimo elemento di un array. Ma se vuoi rimuovere il primo?

Ecco dove entra in gioco `.shift()`. Funziona esattamente come `.pop()`, ma rimuove il primo elemento invece dell'ultimo.

Esempio:

```js
const ourArray = ["Stimpson", "J", ["cat"]];
const removedFromOurArray = ourArray.shift();
```

`removedFromOurArray` avrebbe come valore la stringa `Stimpson`, e `ourArray` avrebbe `["J", ["cat"]]`.

# --instructions--

Usa la funzione `.shift()` per rimuovere il primo elemento da `myArray` e assegna il valore "shifted off" a una nuova variabile, `removedFromMyArray`.

# --hints--

`myArray` dovrebbe ora essere uguale a `[["dog", 3]]`.

```js
assert(
  (function (d) {
    if (d[0][0] == 'dog' && d[0][1] === 3 && d[1] == undefined) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

`removedFromMyArray` dovrebbe contenere `["John", 23]`.

```js
assert(
  (function (d) {
    if (
      d[0] == 'John' &&
      d[1] === 23 &&
      typeof removedFromMyArray === 'object'
    ) {
      return true;
    } else {
      return false;
    }
  })(removedFromMyArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["dog", 3]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["dog", 3]];

// Only change code below this line
const removedFromMyArray = myArray.shift();
```
