---
id: 56bbb991ad1ed5201cd392d0
title: creare oggetti JavaScript
challengeType: 1
forumTopicId: 16769
dashedName: build-javascript-objects
---

# --description--

Potresti aver già sentito il termine `object`.

Gli oggetti sono simili agli `arrays`, ma invece di usare gli indici per consultare e modificare i dati, consulti i dati negli oggetti tramite quelle che si chiamano `properties`.

Gli oggetti sono utili per archiviare dati in modo strutturato e possono rappresentare oggetti del mondo reale, come un gatto.

Ecco un esempio di oggetto gatto:

```js
const cat = {
  "name": "Whiskers",
  "legs": 4,
  "tails": 1,
  "enemies": ["Water", "Dogs"]
};
```

In questo esempio, tutte le proprietà sono archiviate come stringhe, come `name`, `legs` e `tails`. Tuttavia, puoi anche usare numeri come proprietà. Puoi persino omettere le virgolette per le proprietà stringa composte da una sola parola, come segue:

```js
const anotherObject = {
  make: "Ford",
  5: "five",
  "model": "focus"
};
```

Tuttavia, se il tuo oggetto ha proprietà non stringa, JavaScript le convertirà automaticamente in stringhe.

# --instructions--

Crea un oggetto che rappresenti un cane chiamato `myDog` che contenga le proprietà `name` (una stringa), `legs`, `tails` e `friends`.

Puoi impostare queste proprietà dell'oggetto con i valori che vuoi, purché `name` sia una stringa, `legs` e `tails` siano numeri e `friends` sia un array.

# --hints--

`myDog` dovrebbe contenere la proprietà `name` e questa dovrebbe essere una `string`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('name') &&
      z.name !== undefined &&
      typeof z.name === 'string'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` dovrebbe contenere la proprietà `legs` e questa dovrebbe essere una `number`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('legs') &&
      z.legs !== undefined &&
      typeof z.legs === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` dovrebbe contenere la proprietà `tails` e questa dovrebbe essere una `number`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('tails') &&
      z.tails !== undefined &&
      typeof z.tails === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` dovrebbe contenere la proprietà `friends` e questa dovrebbe essere un `array`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('friends') &&
      z.friends !== undefined &&
      Array.isArray(z.friends)
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` dovrebbe contenere solo tutte le proprietà date.

```js
assert(
  (function (z) {
    return Object.keys(z).length === 4;
  })(myDog)
);
```

# --seed--

## --seed-contents--

```js
const myDog = {
  // Only change code below this line


  // Only change code above this line
};
```

# --solutions--

```js
const myDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```
