---
id: 56bbb991ad1ed5201cd392cf
title: Scrivere JavaScript riutilizzabile con le funzioni
challengeType: 1
forumTopicId: 18378
dashedName: write-reusable-javascript-with-functions
---

# --description--

In JavaScript, possiamo suddividere il nostro codice in parti riutilizzabili chiamate <dfn>funzioni</dfn>.

Ecco un esempio di funzione:

```js
function functionName() {
  console.log("Hello World");
}
```

Puoi chiamare o <dfn>invocare</dfn> questa funzione usando il suo nome seguito dalle parentesi tonde, così: `functionName();` Ogni volta che la funzione viene chiamata stamperà il messaggio `Hello World` nella console di sviluppo. Tutto il codice tra le parentesi graffe verrà eseguito ogni volta che la funzione viene chiamata.

# --instructions--

<ol>
  <li>
    Crea una funzione chiamata <code>reusableFunction</code> che stampi la stringa <code>Hi World</code> nella console di sviluppo.
  </li>
  <li>
    Chiama la funzione.
  </li>
</ol>

# --before-each--

```js
function testConsole() {
  var originalConsole = console;
  var nativeLog = console.log;
  var hiWorldWasLogged = false;
  console.log = function (message) {
    if (message === 'Hi World') {
      console.warn(message);
      hiWorldWasLogged = true;
    }
    if (nativeLog.apply) {
      nativeLog.apply(originalConsole, arguments);
    } else {
      var nativeMsg = Array.prototype.slice.apply(arguments).join(' ');
      nativeLog(nativeMsg);
    }
  };
  reusableFunction();
  console.log = nativeLog;
  return hiWorldWasLogged;
}
```

# --hints--

`reusableFunction` dovrebbe essere una funzione.

```js
assert(typeof reusableFunction === 'function');
```

Se `reusableFunction` viene chiamata, dovrebbe stampare la stringa `Hi World` nella console.

```js
assert(testConsole());
```

Dovresti chiamare `reusableFunction` una volta che è definita.

```js
const functionStr = reusableFunction && __helpers.removeWhiteSpace(reusableFunction.toString());
const codeWithoutFunction = __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).replace(/reusableFunction\(\)\{/g, '');
assert(/reusableFunction\(\)/.test(codeWithoutFunction));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reusableFunction() {
  console.log("Hi World");
}
reusableFunction();
```
