---
id: 56bbb991ad1ed5201cd392ca
title: Consultare i dati di un array con gli indici
challengeType: 1
forumTopicId: 16158
dashedName: access-array-data-with-indexes
---

# --description--

Possiamo consultare i dati all'interno degli array usando gli <dfn>indici</dfn>.

Gli indici degli array si scrivono con la stessa notazione a parentesi che usano le stringhe, ma invece di specificare un carattere, indicano una voce nell'array. Come per le stringhe, gli array usano un indice <dfn>a base zero</dfn>, quindi il primo elemento di un array ha indice `0`.

<br>

**Esempio**

```js
const array = [50, 60, 70];
console.log(array[0]);
const data = array[1];
```

Il `console.log(array[0])` stampa `50`, e `data` ha il valore `60`.

# --instructions--

Crea una variabile chiamata `myData` e impostala uguale al primo valore di `myArray` usando la notazione a parentesi.

# --hints--

La variabile `myData` dovrebbe essere uguale al primo valore di `myArray`.

```js
assert(
  (function () {
    if (
      typeof myArray !== 'undefined' &&
      typeof myData !== 'undefined' &&
      myArray[0] === myData
    ) {
      return true;
    } else {
      return false;
    }
  })()
);
```

I dati nella variabile `myArray` devono essere consultati usando la notazione a parentesi.

```js
assert(
  (function () {
    if (__helpers.removeJSComments(code).match(/\s*=\s*myArray\[0\]/g)) {
      return true;
    } else {
      return false;
    }
  })()
);
```

# --seed--

## --seed-contents--

```js
const myArray = [50, 60, 70];


```

# --solutions--

```js
const myArray = [50, 60, 70];
const myData = myArray[0];
```
