---
id: 56bbb991ad1ed5201cd392d3
title: Eliminare proprietà da un oggetto JavaScript
challengeType: 1
forumTopicId: 17560
dashedName: delete-properties-from-a-javascript-object
---

# --description--

Possiamo anche eliminare proprietà dagli oggetti in questo modo:

```js
delete ourDog.bark;
```

Esempio:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"],
  "bark": "bow-wow"
};

delete ourDog.bark;
```

Dopo l’ultima riga mostrata sopra, `ourDog` appare così:

```js
{
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
}
```

# --instructions--

Elimina la proprietà `tails` da `myDog`. Puoi usare sia la notazione a punto sia la notazione a parentesi.

# --hints--

Dovresti eliminare la proprietà `tails` da `myDog`.

```js
assert(typeof myDog === 'object' && myDog.tails === undefined);
```

Non dovresti modificare la configurazione `myDog`.

```js
assert(__helpers.removeJSComments(code).match(/"tails": 1/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"],
  "bark": "woof"
};

// Only change code below this line

```

# --solutions--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"],
  "bark": "woof"
};
delete myDog.tails;
```
