---
id: 56bbb991ad1ed5201cd392d2
title: aggiungere nuove proprietà a un oggetto JavaScript
challengeType: 1
forumTopicId: 301169
dashedName: add-new-properties-to-a-javascript-object
---

# --description--

Puoi aggiungere nuove proprietà agli oggetti JavaScript esistenti nello stesso modo in cui li modifichi.

Ecco come aggiungeremmo una proprietà `bark` a `ourDog`:

```js
ourDog.bark = "bow-wow";
```

oppure

```js
ourDog["bark"] = "bow-wow";
```

Ora, quando valutiamo `ourDog.bark`, otterremo il suo abbaiare, `bow-wow`.

Esempio:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};

ourDog.bark = "bow-wow";
```

# --instructions--

Aggiungi una proprietà `bark` a `myDog` e impostala su un suono di cane, come "woof". Puoi usare sia la notazione a punto sia la notazione a parentesi.

# --hints--

Dovresti aggiungere la proprietà `bark` a `myDog`.

```js
assert(myDog.bark !== undefined);
```

Non dovresti aggiungere `bark` all'inizializzazione di `myDog`.

```js
assert(!/bark[^\n]:/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};


```

# --solutions--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};
myDog.bark = "Woof Woof";
```
