---
id: 56bbb991ad1ed5201cd392d1
title: aggiornare le proprietà di un oggetto
challengeType: 1
forumTopicId: 18336
dashedName: updating-object-properties
---

# --description--

Dopo aver creato un oggetto JavaScript, puoi aggiornare le sue proprietà in qualsiasi momento proprio come faresti con qualsiasi altra variabile. Puoi usare sia la notazione a punto sia la notazione a parentesi per aggiornare.

Per esempio, diamo un’occhiata a `ourDog`:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```

Dato che è un cane particolarmente felice, cambiamo il suo nome nella stringa `Happy Camper`. Ecco come aggiorniamo la proprietà name del suo oggetto: `ourDog.name = "Happy Camper";` oppure `ourDog["name"] = "Happy Camper";` Ora, quando valutiamo `ourDog.name`, invece di ottenere `Camper`, otterremo il suo nuovo nome, `Happy Camper`.

# --instructions--

Aggiorna la proprietà name dell’oggetto `myDog`. Cambiamo il suo nome da `Coder` a `Happy Coder`. Puoi usare sia la notazione a punto sia la notazione a parentesi.

# --hints--

Dovresti aggiornare la proprietà `myDog` di `name` per farla uguale alla stringa `Happy Coder`.

```js
assert(/happy coder/gi.test(myDog.name));
```

Non dovresti modificare la definizione di `myDog`.

```js
assert(/"name": "Coder"/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const myDog = {
  "name": "Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};

// Only change code below this line

```

# --solutions--

```js
const myDog = {
  "name": "Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};
myDog.name = "Happy Coder";
```
