---
id: 587d7b7e367417b2b2512b21
title: Usa operatori condizionali (ternari) multipli
challengeType: 1
forumTopicId: 301179
dashedName: use-multiple-conditional-ternary-operators
---

# --description--

Nella sfida precedente, hai usato un singolo operatore condizionale. Puoi anche concatenarli per verificare più condizioni.

La funzione seguente usa le istruzioni `if`, `else if` e `else` per controllare più condizioni:

```js
function findGreaterOrEqual(a, b) {
  if (a === b) {
    return "a and b are equal";
  }
  else if (a > b) {
    return "a is greater";
  }
  else {
    return "b is greater";
  }
}
```

La funzione sopra può essere riscritta usando operatori condizionali multipli:

```js
function findGreaterOrEqual(a, b) {
  return (a === b) ? "a and b are equal" 
    : (a > b) ? "a is greater" 
    : "b is greater";
}
```

È considerata una buona pratica formattare gli operatori condizionali multipli in modo che ogni condizione sia su una riga separata, come mostrato sopra. Usare operatori condizionali multipli senza una corretta indentazione può rendere il tuo codice difficile da leggere. Per esempio:

```js
function findGreaterOrEqual(a, b) {
  return (a === b) ? "a and b are equal" : (a > b) ? "a is greater" : "b is greater";
}
```

# --instructions--

Nella funzione `checkSign`, usa operatori condizionali multipli — seguendo il formato consigliato usato in `findGreaterOrEqual` — per verificare se un numero è positivo, negativo o zero. La funzione dovrebbe restituire `positive`, `negative` o `zero`.

# --hints--

`checkSign` dovrebbe usare operatori condizionali multipli

```js
assert(/.+?\s*?\?\s*?.+?\s*?:\s*?.+?\s*?\?\s*?.+?\s*?:\s*?.+?/gi.test(__helpers.removeJSComments(code)));
```

`checkSign(10)` dovrebbe restituire la stringa `positive`. Nota che la maiuscola/minuscola conta

```js
assert(checkSign(10) === 'positive');
```

`checkSign(-12)` dovrebbe restituire la stringa `negative`. Nota che la maiuscola/minuscola conta

```js
assert(checkSign(-12) === 'negative');
```

`checkSign(0)` dovrebbe restituire la stringa `zero`. Nota che la maiuscola/minuscola conta

```js
assert(checkSign(0) === 'zero');
```

# --seed--

## --seed-contents--

```js
function checkSign(num) {

}

checkSign(10);
```

# --solutions--

```js
function checkSign(num) {
  return (num > 0) ? 'positive' : (num < 0) ? 'negative' : 'zero';
}
```
