---
id: 587d7b7e367417b2b2512b23
title: Usa la funzione parseInt
challengeType: 1
forumTopicId: 301183
dashedName: use-the-parseint-function
---

# --description--

La funzione `parseInt()` analizza una stringa e restituisce un intero. Ecco un esempio:

```js
const a = parseInt("007");
```

La funzione sopra converte la stringa `007` nell'intero `7`. Se il primo carattere nella stringa non può essere convertito in un numero, allora restituisce `NaN`.

# --instructions--

Usa `parseInt()` nella funzione `convertToInteger` in modo che converta la stringa di input `str` in un intero e la restituisca.

# --hints--

`convertToInteger` dovrebbe usare la funzione `parseInt()`

```js
assert(/parseInt/g.test(__helpers.removeJSComments(code)));
```

`convertToInteger("56")` dovrebbe restituire un numero

```js
assert(typeof convertToInteger('56') === 'number');
```

`convertToInteger("56")` dovrebbe restituire 56

```js
assert(convertToInteger('56') === 56);
```

`convertToInteger("77")` dovrebbe restituire 77

```js
assert(convertToInteger('77') === 77);
```

`convertToInteger("JamesBond")` dovrebbe restituire `NaN`

```js
assert.isNaN(convertToInteger('JamesBond'));
```

# --seed--

## --seed-contents--

```js
function convertToInteger(str) {

}

convertToInteger("56");
```

# --solutions--

```js
function convertToInteger(str) {
  return parseInt(str);
}
```
