---
id: 587d7b7e367417b2b2512b22
title: Usa la funzione parseInt con una base
challengeType: 1
forumTopicId: 301182
dashedName: use-the-parseint-function-with-a-radix
---

# --description--

La funzione `parseInt()` analizza una stringa e restituisce un intero. Prende un secondo argomento per la base, che specifica la base del numero nella stringa. La base può essere un intero compreso tra 2 e 36.

La chiamata alla funzione è così:

```js
parseInt(string, radix);
```

Ecco un esempio:

```js
const a = parseInt("11", 2);
```

La variabile base indica che `11` è nel sistema binario, o base 2. Questo esempio converte la stringa `11` in un intero `3`.

# --instructions--

Usa `parseInt()` nella funzione `convertToInteger` in modo che converta un numero binario in un intero e lo restituisca.

# --hints--

`convertToInteger` dovrebbe usare la funzione `parseInt()`

```js
assert(/parseInt/g.test(__helpers.removeJSComments(code)));
```

`convertToInteger("10011")` dovrebbe restituire un numero

```js
assert(typeof convertToInteger('10011') === 'number');
```

`convertToInteger("10011")` dovrebbe restituire 19

```js
assert(convertToInteger('10011') === 19);
```

`convertToInteger("111001")` dovrebbe restituire 57

```js
assert(convertToInteger('111001') === 57);
```

`convertToInteger("JamesBond")` dovrebbe restituire `NaN`

```js
assert.isNaN(convertToInteger('JamesBond'));
```

# --seed--

## --seed-contents--

```js
function convertToInteger(str) {

}

convertToInteger("10011");
```

# --solutions--

```js
function convertToInteger(str) {
  return parseInt(str, 2);
}
```
