---
id: 587d7b87367417b2b2512b3f
title: Esplora le differenze tra le parole chiave var e let
challengeType: 1
forumTopicId: 301202
dashedName: explore-differences-between-the-var-and-let-keywords
---

# --description--

Uno dei problemi più grandi nel dichiarare variabili con la parola chiave `var` è che puoi facilmente sovrascrivere dichiarazioni di variabili:

```js
var camper = "James";
var camper = "David";
console.log(camper);
```

Nel codice sopra, la variabile `camper` è inizialmente dichiarata come `James`, e poi viene sovrascritta diventando `David`. La console mostra quindi la stringa `David`.

In una piccola applicazione, potresti non incontrare questo tipo di problema. Ma man mano che il tuo codice diventa più grande, potresti accidentalmente sovrascrivere una variabile che non intendevi modificare. Poiché questo comportamento non genera un errore, cercare e correggere bug diventa più difficile.

Una parola chiave chiamata `let` è stata introdotta in ES6, un importante aggiornamento di JavaScript, per risolvere questo potenziale problema con la parola chiave `var`. Imparerai altre funzionalità di ES6 in sfide successive.

Se sostituisci `var` con `let` nel codice sopra, si verifica un errore:

```js
let camper = "James";
let camper = "David";
```

L'errore può essere visto nella console del browser.

Quindi, a differenza di `var`, quando usi `let`, una variabile con lo stesso nome può essere dichiarata solo una volta.

# --instructions--

Aggiorna il codice in modo che usi solo la parola chiave `let`.

# --hints--

`var` non dovrebbe esistere nel codice.

```js
assert.notMatch(code, /var/g);
```

`catName` dovrebbe essere la stringa `Oliver`.

```js
assert.equal(catName, 'Oliver');
```

`catSound` dovrebbe essere la stringa `Meow!`

```js
assert.equal(catSound, 'Meow!');
```

# --seed--

## --seed-contents--

```js
var catName = "Oliver";
var catSound = "Meow!";
```

# --solutions--

```js
let catName = "Oliver";
let catSound = "Meow!";
```
