---
id: 5cfa3679138e7d9595b9d9d4
title: sostituire i loop usando la ricorsione
challengeType: 1
videoUrl: >-
  https://www.freecodecamp.org/news/how-recursion-works-explained-with-flowcharts-and-a-video-de61f40cb7f9/
forumTopicId: 301175
dashedName: replace-loops-using-recursion
---

# --description--

La ricorsione è il concetto per cui una funzione può essere espressa in termini di se stessa. Per aiutarti a capire, inizia pensando alla seguente attività: moltiplicare i primi `n` elementi di un array per ottenere il prodotto di quegli elementi. Usando un loop `for`, potresti fare così:

```js
  function multiply(arr, n) {
    let product = 1;
    for (let i = 0; i < n; i++) {
      product *= arr[i];
    }
    return product;
  }
```

Tuttavia, nota che `multiply(arr, n) == multiply(arr, n - 1) * arr[n - 1]`. Questo significa che puoi riscrivere `multiply` in termini di se stessa senza mai dover usare un loop.

```js
  function multiply(arr, n) {
    if (n <= 0) {
      return 1;
    } else {
      return multiply(arr, n - 1) * arr[n - 1];
    }
  }
```

La versione ricorsiva di `multiply` si scompone così. Nel <dfn>caso base</dfn>, dove `n <= 0`, restituisce 1. Per valori più grandi di `n`, si chiama da sola, ma con `n - 1`. Questa chiamata di funzione viene valutata allo stesso modo, chiamando `multiply` di nuovo finché `n <= 0`. A questo punto, tutte le funzioni possono restituire e il `multiply` originale restituisce la risposta.

**Nota:** Le funzioni ricorsive devono avere un caso base in cui restituiscono senza chiamare di nuovo la funzione (in questo esempio, quando `n <= 0`), altrimenti non possono mai terminare l'esecuzione.

# --instructions--

Scrivi una funzione ricorsiva, `sum(arr, n)`, che restituisca la somma dei primi `n` elementi di un array `arr`.

# --hints--

`sum([1], 0)` dovrebbe essere uguale a 0.

```js
assert.equal(sum([1], 0), 0);
```

`sum([2, 3, 4], 1)` dovrebbe essere uguale a 2.

```js
assert.equal(sum([2, 3, 4], 1), 2);
```

`sum([2, 3, 4, 5], 3)` dovrebbe essere uguale a 9.

```js
assert.equal(sum([2, 3, 4, 5], 3), 9);
```

Il tuo codice non dovrebbe fare affidamento su alcun tipo di loop (`for` o `while` o funzioni di ordine superiore come `forEach`, `map`, `filter` o `reduce`).

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

Dovresti usare la ricorsione per risolvere questo problema.

```js
assert(
  sum.toString().match(/sum\(.*\)/g).length > 1
);
```

# --seed--

## --seed-contents--

```js
function sum(arr, n) {
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function sum(arr, n) {
  // Only change code below this line
  if(n <= 0) {
    return 0;
  } else {
    return sum(arr, n - 1) + arr[n - 1];
  }
  // Only change code above this line
}
```
