---
id: 5ee127a03c3b35dd45426493
title: "Assegnare il valore di una variabile a un'altra"
challengeType: 1
forumTopicId: 418265
dashedName: assigning-the-value-of-one-variable-to-another
---

# --description--

Dopo che un valore è stato assegnato a una variabile usando l’operatore di <dfn>assegnazione</dfn>, puoi assegnare il valore di quella variabile a un’altra variabile usando l’operatore di <dfn>assegnazione</dfn>.

```js
var myVar;
myVar = 5;
var myNum;
myNum = myVar;
```

La riga sopra dichiara una variabile `myVar` senza valore, poi le assegna il valore `5`. Successivamente, viene dichiarata una variabile chiamata `myNum` senza valore. Poi, il contenuto di `myVar` (che è `5`) viene assegnato alla variabile `myNum`. Ora, `myNum` ha anch’essa il valore di `5`.

# --instructions--

Assegna il contenuto di `a` alla variabile `b`.

# --hints--

Non dovresti modificare il codice sopra il commento specificato.

```js
assert(/var a;/.test(__helpers.removeJSComments(code)) && /a = 7;/.test(__helpers.removeJSComments(code)) && /var b;/.test(__helpers.removeJSComments(code)));
```

`b` dovrebbe avere il valore `7`.

```js
assert(typeof b === 'number' && b === 7);
```

`a` dovrebbe essere assegnato a `b` con `=`.

```js
assert(/b\s*=\s*a\s*/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
var a;
a = 7;
var b;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
var b;
b = a;
```
