---
id: bd7993c9c69feddfaeb7bdef
title: Moltiplicare due numeri decimali con JavaScript
challengeType: 1
forumTopicId: 301173
dashedName: multiply-two-decimals-with-javascript
---

# --description--

In JavaScript puoi anche fare calcoli con numeri decimali, proprio come con i numeri interi.

Moltiplichiamo due numeri decimali per ottenere il loro prodotto.

# --instructions--

Modifica `0.0` in modo che product sia uguale a `5.0`.

# --hints--

La variabile `product` dovrebbe essere uguale a `5.0`.

```js
assert(product === 5.0);
```

Dovresti usare l'operatore `*`

```js
assert(/\*/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const product = 2.0 * 0.0;
```

# --solutions--

```js
const product = 2.0 * 2.5;
```
