---
id: cf1111c1c11feddfaeb3bdef
title: Aggiungere due numeri con JavaScript
challengeType: 1
forumTopicId: 16650
dashedName: add-two-numbers-with-javascript
---

# --description--

`Number` è un tipo di dato in JavaScript che rappresenta dati numerici.

Ora proviamo ad aggiungere due numeri usando JavaScript.

JavaScript usa il simbolo `+` come operatore di addizione quando è posto tra due numeri.

**Esempio:**

```js
const myVar = 5 + 10;
```

`myVar` ora ha il valore `15`.

# --instructions--

Modifica `0` in modo che la somma sia uguale a `20`.

# --hints--

`sum` dovrebbe essere uguale a `20`.

```js
assert(sum === 20);
```

Dovresti usare l'operatore `+`.

```js
assert(/\+/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const sum = 10 + 0;
```

# --solutions--

```js
const sum = 10 + 10;
```
