---
id: cf1231c1c11feddfaeb5bdef
title: Moltiplicare due numeri con JavaScript
challengeType: 1
forumTopicId: 18243
dashedName: multiply-two-numbers-with-javascript
---

# --description--

Possiamo anche moltiplicare un numero per un altro.

JavaScript usa il simbolo `*` per la moltiplicazione di due numeri.

**Esempio**

```js
const myVar = 13 * 13;
```

`myVar` avrebbe il valore `169`.

# --instructions--

Modifica `0` in modo che il prodotto sia uguale a `80`.

# --hints--

La variabile `product` dovrebbe essere uguale a 80.

```js
assert(product === 80);
```

Dovresti usare l'operatore `*`.

```js
assert(/\*/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const product = 8 * 0;
```

# --solutions--

```js
const product = 8 * 10;
```
