---
id: cf1391c1c11feddfaeb4bdef
title: Creare numeri decimali con JavaScript
challengeType: 1
forumTopicId: 16826
dashedName: create-decimal-numbers-with-javascript
---

# --description--

Possiamo archiviare numeri decimali anche nelle variabili. I numeri decimali sono talvolta chiamati numeri <dfn>floating point</dfn> o <dfn>float</dfn>.

**Nota:** quando calcoli numeri, questi vengono calcolati con precisione finita. Le operazioni con i floating point possono portare a risultati diversi da quelli desiderati. Se ottieni uno di questi risultati, apri un argomento sul <a href="https://forum.freecodecamp.org/" target="_blank" rel="noopener noreferrer nofollow">forum di freeCodeCamp</a>.

# --instructions--

Crea una variabile `myDecimal` e assegnale un valore decimale con una parte frazionaria (ad esempio `5.7`).

# --hints--

`myDecimal` dovrebbe essere un numero.

```js
assert(typeof myDecimal === 'number');
```

`myDecimal` dovrebbe avere un punto decimale

```js
assert(myDecimal % 1 != 0);
```

# --seed--

## --seed-contents--

```js
const ourDecimal = 5.7;

// Only change code below this line

```

# --solutions--

```js
const myDecimal = 9.9;
```
