---
id: 66ed41f912d0bb1dc62da5dd
title: JavaScript 퀴즈 봇 만들기
challengeType: 26
dashedName: lab-javascript-trivia-bot
---

# --description--

**목표:** 아래 사용자 스토리를 완성하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 사용자에게 인사 메시지로 `"Hello! I'm your coding fun fact guide!"`를 콘솔에 출력해야 합니다.
1. 봇의 이름, 출신지, 가장 좋아하는 코딩 언어를 각각 저장하는 `botName`, `botLocation`, `favoriteLanguage` 세 변수를 만들어야 합니다.
1. `"My name is (botName) and I live on (botLocation)."`를 콘솔에 출력해야 합니다.
1. `"My favorite programming language is (favoriteLanguage)."`를 콘솔에 출력해야 합니다.
1. `let`를 사용해 `codingFact` 변수를 만들고, 봇이 가장 좋아하는 코딩 언어에 관한 재미있는 사실을 문자열로 할당하며 `favoriteLanguage` 변수를 포함해야 합니다.
1. `codingFact`를 콘솔에 출력해야 합니다.
1. `codingFact` 변수에 봇이 가장 좋아하는 언어에 관한 새로운 사실을 `favoriteLanguage` 변수를 사용해 다시 할당해야 합니다.
1. `codingFact`를 다시 콘솔에 출력해야 합니다.
1. `codingFact` 변수에 또 다른 새로운 사실을 `favoriteLanguage` 변수를 사용해 다시 할당해야 합니다.
1. `codingFact`를 세 번째로 콘솔에 출력해야 합니다.
1. 봇의 작별 인사로 `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."`를 콘솔에 출력해야 합니다.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

콘솔에 `"Hello! I'm your coding fun fact guide!"`를 출력하세요.

```js
assert.equal(getLogs()[0], "Hello! I'm your coding fun fact guide!")
```

`botName` 변수를 선언하세요. 철자나 대소문자 오류가 없는지 다시 확인하세요.

```js
assert.exists(botName);
```

`botName` 변수는 문자열이어야 합니다.

```js
assert.isString(botName);
```

`botLocation` 변수를 선언하세요. 철자나 대소문자 오류가 없는지 다시 확인하세요.

```js
assert.exists(botLocation);
```

`botLocation` 변수는 문자열이어야 합니다.

```js
assert.isString(botLocation);
```

`favoriteLanguage` 변수를 선언하세요. 철자나 대소문자 오류가 없는지 다시 확인하세요.

```js
assert.exists(favoriteLanguage);
```

`favoriteLanguage` 변수는 문자열이어야 합니다.

```js
assert.isString(favoriteLanguage);
```

콘솔에 `"My name is (botName) and I live on (botLocation)."`를 출력하고 변수들을 문자열에 추가하세요.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[1], `My name is ${botName} and I live on ${botLocation}.`)
```

콘솔에 `"My favorite programming language is (favoriteLanguage)."`를 출력하고 변수를 문자열에 추가하세요.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[2], `My favorite programming language is ${favoriteLanguage}.`)
```

`let`를 사용해 새 변수 `codingFact`를 선언하세요.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/)
```

`codingFact`에 `favoriteLanguage`를 포함하는 값을 할당하세요.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/);
assert.match(code, /codingFact\s*=\s*(("|'|`)?.+?\1?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|'|`)?.+?\3?))/);
assert.exists(codingFact);
assert.isNotEmpty(codingFact);
assert.include(String(codingFact), favoriteLanguage);
assert.include(getLogs()[3], favoriteLanguage);
```

콘솔에 `codingFact`를 출력하세요.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[3], favoriteLanguage);
assert.isAtLeast(loggingCodingFacts.length, 1);
```

`codingFact`에 `favoriteLanguage`를 포함하는 새 값을 할당하고 콘솔에 출력하세요.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
const [first, second, third] = codeWithoutComments.match(/(let )?\s*codingFact\s*=\s*(("|')?.+?\2?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|')?.+?\2?))/g);
assert.include(getLogs()[4], favoriteLanguage);
assert.notEqual(getLogs()[4], getLogs()[3]);
assert.isAtLeast(loggingCodingFacts.length, 2);
assert.exists(second);
assert.isNotEmpty(codingFact); 
```

`codingFact`에 `favoriteLanguage`를 포함하는 값을 세 번째로 할당하고 콘솔에 출력하세요.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[5], favoriteLanguage);
assert.notEqual(getLogs()[5], getLogs()[4]);
assert.equal(getLogs()[5], codingFact);
assert.lengthOf(loggingCodingFacts, 3);
assert.isNotEmpty(codingFact); 
```

콘솔에 `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."`를 출력하고 변수들의 값을 추가하세요.

```js
assert.equal(getLogs()[6], `It was fun sharing these facts with you. Goodbye! - ${botName} from ${botLocation}.`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
console.log("Hello! I'm your coding fun fact guide!");

const botName = "JsBot";
const botLocation = "planet Eris";
const favoriteLanguage = "JavaScript";

console.log("My name is " + botName + " and I live on " + botLocation + ".");
console.log("My favorite programming language is " + favoriteLanguage + ".");

let codingFact = "Did you know that " + favoriteLanguage + " was created in just 10 days?";

console.log(codingFact);

codingFact = "Another fun fact: " + favoriteLanguage + "was originally called Mocha!";

console.log(codingFact);

codingFact = "Also, " + favoriteLanguage + " is the most popular programming language in the world.";

console.log(codingFact);

console.log("It was fun sharing these facts with you. Goodbye! - " + botName + " from " + botLocation + ".");

```
