---
id: 619665c9abd72906f3ad30f9
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

ستبني بطريق Flappy سعيد، وستستكشف المزيد حول تحولات CSS والرسوم المتحركة خلال العملية.

ابدأ بقالب HTML الأساسي. أدرج إعلان `DOCTYPE`، وعنصر `html` مع تعيين اللغة إلى الإنجليزية، ووسوم `meta` المناسبة، وعنصر `head`، و`body`، و`title`. أيضًا، اربط ورقة الأنماط بالصفحة.

# --hints--

يجب أن يحتوي شفرتك على إعلان `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

يجب أن يحتوي شفرتك على عنصر `html`.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

يجب أن يحتوي عنصر `html` على وسم بداية مع خاصية `lang` بقيمة `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصر `html` على وسم الإغلاق.

```js
assert(code.match(/<\/html\s*>/gi));
```

يجب أن يحتوي شفرتك على عنصر `head` داخل عنصر `html`.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

يجب أن يحتوي شفرتك على عنصر `body` داخل عنصر `html`.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

يجب أن يأتي عنصر `head` قبل عنصر `body`.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

يجب أن يكون لديك وسمان `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

يجب أن يحتوي أحد وسوم `meta` على خاصية `name` بقيمة `viewport`، و`content` بقيمة `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

يجب أن يحتوي وسم `meta` الآخر على خاصية `charset` بقيمة `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

يجب أن يحتوي شفرتك على عنصر `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

يجب أن يحتوي عنصر `title` على بعض النص.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

يجب أن يحتوي شفرتك على عنصر `link`.

```js
assert.exists(document.querySelector('link'));
```

يجب أن يكون عنصر `link` داخل عنصر `head`.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i));
```

يجب أن يحتوي عنصر `link` على خاصية `rel` بقيمة `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

يجب أن يحتوي عنصر `link` على خاصية `href` بقيمة `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
