---
id: 662a6bc12cde72c32fb526f0
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

Ein <dfn>Interface</dfn> ist wie ein Bauplan für eine Klasse. Ein Interface enthält eine Reihe von Methoden und Eigenschaften, die eine Klasse implementieren sollte.

Beginnen Sie dieses Projekt, indem Sie eine leere Klasse mit dem Namen `Equation` deklarieren. Sie verwenden diese Klasse, um ein Interface zu definieren, einen Bauplan für eine generische Gleichung.

# --hints--

Sie sollten eine neue Klasse mit dem Namen `Equation` definieren.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
