---
id: 662bd552e1c1d2db1b88ba47
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Definieren Sie nun eine weitere Klasse namens `LinearEquation` und lassen Sie sie von `Equation` erben. Sie verwenden diese Klasse, um lineare Gleichungen darzustellen.

# --hints--

Sie sollten eine Klasse namens `LinearEquation` definieren.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("LinearEquation")`)) })
```

Ihre `LinearEquation`-Klasse sollte von der `Equation`-Klasse erben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").inherits_from("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
--fcc-editable-region--

--fcc-editable-region--
```
