---
id: 662bdd364bf2cde1487922a9
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Im Gegensatz zu anderen Programmiersprachen implementiert Python keine Interfaces in seiner Kernsprache, aber die Python-Standardbibliothek ermöglicht es Ihnen, Interfaces auf einfache Weise zu definieren.

Für dieses Projekt verwenden Sie Hilfsfunktionen aus dem `abc`-Modul. Importieren Sie daher dieses Modul in Ihren Code.

# --hints--

Sie sollten das `abc`-Modul importieren.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("import abc")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
    

eq = Equation()
lin_eq = LinearEquation()
```
