---
id: 662bde88dc84f1e249801b1a
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

`ABC` steht für *Abstract Base Classes*. Die `ABC`-Klasse ermöglicht es Ihnen, eine normale Klasse in eine abstrakte Klasse zu verwandeln, also eine Klasse, die als Vorlage für konkrete Klassen dient.

Ändern Sie Ihre `import`-Anweisung so, dass nur die `ABC`-Klasse aus dem `abc`-Modul importiert wird. Sie können ein bestimmtes Objekt `x` aus einem Modul `y` mit der Import-Anweisung `from y import x` importieren.

Machen Sie dann Ihre `Equation`-Klasse zu einer abstrakten Klasse, indem Sie sie von `ABC` erben lassen.

# --hints--

Sie sollten `ABC` aus dem `abc`-Modul importieren.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("from abc import ABC")`)) })
```

Ihre `Equation`-Klasse sollte von `ABC` erben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").inherits_from("ABC")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import abc


class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass
    
eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
