---
id: 662bd8260da84bdd5feae419
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Sie möchten, dass die `LinearEquation`-Klasse alle in der `Equation`-Klasse definierten Methoden implementiert und nicht einfach erbt, da diese als Interface fungieren soll.

Derzeit ist die `Equation`-Klasse lediglich die Elternklasse von `LinearEquation`. In den nächsten Schritten lernen Sie, wie Sie sie in ein formales Interface umwandeln.

Erstellen Sie vorerst eine Instanz von `Equation` und weisen Sie sie einer Variablen `eq` zu, sowie eine Instanz von `LinearEquation`, die Sie einer Variablen `lin_eq` zuweisen.

# --hints--

Sie sollten eine Variable `eq` deklarieren und ihr eine Instanz von `Equation` zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("eq = Equation()")`)) })
```

Sie sollten eine Variable `lin_eq` deklarieren und ihr eine Instanz von `LinearEquation` zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation()")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
--fcc-editable-region--

--fcc-editable-region--
```
