---
id: 662bd456896f16d9bd03f1a6
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Definieren Sie innerhalb der `Equation`-Klasse zwei neue Instanzmethoden mit den Namen `solve` und `analyze`.

# --hints--

Sie sollten eine Methode mit dem Namen `solve` innerhalb der `Equation`-Klasse definieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("solve")`)) })
```

Ihre `solve`-Methode sollte einen Parameter namens `self` annehmen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_args("self")`)) })
```

Sie sollten eine Methode mit dem Namen `analyze` innerhalb der `Equation`-Klasse definieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("analyze")`)) })
```

Ihre `analyze`-Methode sollte einen Parameter namens `self` annehmen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Equation:
    pass
--fcc-editable-region--
```
