---
id: 662f6d7c92381a3049e4c987
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Ein Interface muss nicht nur abstrakte Methoden definieren, sondern kann auch Methoden implementieren, die von den konkreten Klassen geerbt werden.

Bevor Sie sich um die eigentliche Implementierung von `solve` und `analyze` in der `Equation`-Klasse kümmern, definieren Sie eine `__init__`-Methode. Verwenden Sie keinen Dekorator dafür.

# --hints--

Sie sollten eine `__init__`-Methode in Ihrer `Equation`-Klasse definieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__init__")`)) })
```

Ihre `__init__`-Methode sollte einen Parameter namens `self` erhalten.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC, abstractmethod


class Equation(ABC):
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
--fcc-editable-region--
```
