---
id: 662fa2e2cf27c09f21f4f5d0
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

In Python werden Datentypen zur Laufzeit erkannt (wenn der Code ausgeführt wird). Daher müssen Sie den Datentyp einer Variablen bei der Deklaration nicht angeben. Dennoch können Sie eine Variable annotieren, um klarzustellen, dass sie einen bestimmten Datentyp mit `variable: <data type> = value` oder einfach `variable: <data type>` enthalten wird. Beachten Sie, dass der Python-Interpreter die zur Annotation verwendeten Typen nicht erzwingt, und normalerweise benötigen Sie externe Werkzeuge dafür.

Definieren Sie innerhalb der `Equation`-Klasse ein Klassenattribut `degree`. Weisen Sie ihm keinen Wert zu. Verwenden Sie stattdessen eine Typannotation von `int`, um zu zeigen, dass es in den konkreten Klassen eine ganze Zahl speichern wird.

Später verwenden Sie dieses Klassenattribut als Teil des Validierungsprozesses der Argumente, die zur Instanziierung der Gleichungsobjekte übergeben werden.

# --hints--

Sie sollten ein Klassenattribut namens `degree` definieren und es mit `int` innerhalb der `Equation`-Klasse annotieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_variable("degree").is_equivalent("degree: int")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    def __init__(self):
        pass
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass
--fcc-editable-region--
lin_eq = LinearEquation()
```
