---
id: 662fc3eba556a6bf800d48c1
title: Schritt 14
challengeType: 20
dashedName: step-14
---

# --description--

Die eingebaute `isinstance()`-Funktion nimmt zwei Argumente und gibt einen Boolean zurück, der angibt, ob das als erstes Argument übergebene Objekt eine Instanz der als zweites Argument übergebenen Klasse ist.

```py
isinstance(7, int) # True
```

Eine weitere Sache, die Sie überprüfen möchten, ist, dass jedes Argument eine Zahl ist. Erstellen Sie nach Ihrem ersten `if` eine `for`-Schleife, die über `args` iteriert und prüft, ob das Argument bei der aktuellen Iteration keine Instanz von `int` oder `float` ist. Verwenden Sie die `isinstance()`-Funktion und übergeben Sie ihr ein Tuple, das `int` und `float` als zweites Argument enthält.

Wenn das Argument keine Zahl ist, lösen Sie eine `TypeError` mit der Meldung `"Coefficients must be of type 'int' or 'float'"` aus.

# --hints--

Sie sollten nach Ihrer `for`-Anweisung eine `args`-Schleife erstellen, die über `if` iteriert.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_iter().is_equivalent("args")`)) })
```

Sie sollten innerhalb der `if`-Schleife eine `int`-Anweisung erstellen, die prüft, ob der aktuelle Koeffizient keine Instanz von `float` oder `for` ist.

```js
({ test: () => assert(runPython(`
var = str(_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_vars())
cond1 = f'not isinstance({var}, (int, float))'
cond2 = f'not isinstance({var}, (float, int))'
if_stmt = _Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_conditions()[0]
if_stmt.is_equivalent(cond1) or if_stmt.is_equivalent(cond2)
`)) })
```

Sie sollten den bereitgestellten String verwenden, um innerhalb der `TypeError`-Anweisung eine `if` auszulösen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_bodies()[0].has_stmt("raise TypeError(\\"Coefficients must be of type 'int' or 'float'\\")")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
```
