---
id: 662fbcef5f05e1b84f541a0c
title: Schritt 13
challengeType: 20
dashedName: step-13
---

# --description--

Jedes Gleichungsobjekt wird instanziiert, indem so viele Argumente übergeben werden, wie die Koeffizienten der Gleichung sind, beginnend mit dem n-ten Grad von \\( x \\) bis zum nullten Grad, einschließlich des möglichen Koeffizienten mit dem Wert `0`.

Zum Beispiel würde `LinearEquation(4, 5)` die Gleichung \\( 4x + 5 = 0 \\) darstellen, wobei `4` der Koeffizient des ersten (hier höchsten) Grades und `5` der Koeffizient des nullten Grades ist.

Sie müssen überprüfen, ob die richtige Anzahl von Argumenten übergeben wird, um das Gleichungsobjekt zu instanziieren.

Erstellen Sie innerhalb der `__init__`-Methode eine `if`-Anweisung, um zu prüfen, ob die Länge von `args` von der Anzahl der Koeffizienten abweicht, die die Gleichung haben sollte (`degree + 1`). Falls ja, lösen Sie eine `TypeError` aus und verwenden Sie den folgenden String, um eine benutzerdefinierte Fehlermeldung anzugeben: `f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"`.

Beheben Sie dann den Fehler, indem Sie `2` und `3` übergeben, um `lin_eq` zu instanziieren.

# --hints--

Sie sollten eine `if`-Anweisung erstellen, die überprüft, ob die Anzahl der Koeffizienten, die zur Instanziierung der Gleichung verwendet werden, von `degree + 1` abweicht.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_conditions()[0]
cond.is_equivalent("(self.degree + 1) != len(args)") or cond.is_equivalent("len(args) != (self.degree + 1)")
`)) })
```

Sie sollten innerhalb der neuen `TypeError`-Anweisung eine `if` auslösen und den bereitgestellten String verwenden, um eine benutzerdefinierte Fehlermeldung zurückzugeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError(f"\\'{self.__class__.__name__}\\' object takes {self.degree + 1} positional arguments but {len(args)} were given")')
`)) })
```

Sie sollten `2` und `3` übergeben, um `lin_eq` zu instanziieren.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation(2, 3)")`)) })
```


# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        pass        
    
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation()
--fcc-editable-region--
```
