---
id: 662f96576ef178927de87975
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Damit eine Methode als abstrakte Methode erkannt wird, sollte sie mit dem `@abstractmethod`-Dekorator versehen sein.

Ein <dfn>Dekorator</dfn> wird in Python verwendet, um das Verhalten einer Funktion zu verändern. Hier ist ein Beispiel, wie man einen Dekorator namens `spam` verwendet:

```py
@spam
def foo():
    pass
```

Ändern Sie Ihre Import-Anweisung so, dass der `abstractmethod`-Dekorator importiert wird, und dekorieren Sie sowohl die `solve`- als auch die `analyze`-Methoden der `Equation`-Klasse. Dadurch werden zwei Ausnahmen ausgelöst.

Sobald eine von `ABC` abgeleitete Klasse eine abstrakte Methode besitzt, kann die Klasse nicht mehr instanziiert werden. Löschen Sie daher die `Equation`-Instanz, um den Fehler zu beheben.

Der andere Fehler tritt auf, weil die `LinearEquation`-Klasse alle abstrakten Methoden, die im Interface definiert sind, implementieren muss. Stellen Sie sicher, dass Sie diese auch in der `LinearEquation`-Klasse definieren. Sie dürfen den `abstractmethod`-Dekorator in der konkreten Klasse nicht verwenden.

# --hints--

Sie sollten `abstractmethod` aus dem `abc`-Modul importieren.

```js
({ test: () => assert(runPython(`
_Node(_code).has_import("from abc import ABC, abstractmethod") or \\
_Node(_code).has_import("from abc import abstractmethod, ABC") or \\
(_Node(_code).has_import("from abc import abstractmethod") and _Node(_code).has_import("from abc import ABC"))
`)) })
```

Sie sollten die `@abstractmethod`-Methode innerhalb der `solve`-Klasse mit `Equation` dekorieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_decorators("abstractmethod")`)) })
```

Sie sollten die `@abstractmethod`-Methode innerhalb der `analyze`-Klasse mit `Equation` dekorieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_decorators("abstractmethod")`)) })
```

Sie sollten eine Methode namens `solve` innerhalb der `LinearEquation`-Klasse definieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("solve")`)) })
```

Ihre `solve`-Methode sollte einen Parameter namens `self` annehmen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_args("self")`)) })
```

Sie sollten eine Methode namens `analyze` innerhalb der `LinearEquation`-Klasse definieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("analyze")`)) })
```

Ihre `solve`-Methode sollte einen Parameter namens `self` annehmen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC


class Equation(ABC):
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass

eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
