---
id: 6639f947d3a1818c9322c64a
title: Schritt 16
challengeType: 20
dashedName: step-16
---

# --description--

Der letzte Schritt bei der Validierung der Koeffizienten besteht darin, zu überprüfen, dass der Koeffizient des höchsten Grades ungleich null ist. Denken Sie daran, dass der Koeffizient des höchsten Grades als erstes Argument beim Instanziieren des Objekts übergeben werden sollte.

Fügen Sie dafür eine `if`-Anweisung hinzu und werfen Sie eine `ValueError` mit dem folgenden String, um eine benutzerdefinierte Nachricht auszugeben: `'Highest degree coefficient must be different from zero'`.

# --hints--

Sie sollten eine `if`-Anweisung erstellen, die überprüft, ob der erste übergebene Koeffizient zum Instanziieren der Gleichung gleich null ist.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[2].find_conditions()[0]
cond.is_equivalent("args[0] == 0") or cond.is_equivalent("0 == args[0]") or cond.is_equivalent("not args[0]")
`)) })
```

Sie sollten innerhalb der neuen `ValueError`-Anweisung eine `if` auslösen und den bereitgestellten String verwenden, um eine benutzerdefinierte Fehlermeldung zurückzugeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[2].find_bodies()[0].has_stmt("raise ValueError('Highest degree coefficient must be different from zero')")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
```
