---
id: 663a22ba7420c4d2f7fd2aec
title: Schritt 25
challengeType: 20
dashedName: step-25
---

# --description--

Es ist Zeit, die `solve`-Methode zu implementieren. Gegeben ist eine lineare Gleichung in der Form \\( ax + b = 0 \\), deren Lösung \\(x = -\frac{b}{a}\\) lautet.

Entpacken Sie die Koeffizienten, die im `coefficients`-Attribut gespeichert sind, in die Variablen `a` und `b`. Beachten Sie, dass Sie die `.values()`-Methode verwenden müssen.

Deklarieren Sie dann eine Variable `x`, weisen Sie ihr die Lösung der Gleichung zu und geben Sie sie aus der `solve`-Methode zurück.

# --hints--

Sie sollten die Werte, die im `coefficients`-Attribut gespeichert sind, in die Variablen `a` und `b` entpacken.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_stmt("a, b = self.coefficients.values()")`)) })
```

Sie sollten eine Variable mit dem Namen `x` deklarieren und ihr die Lösung der linearen Gleichung zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_stmt("x = -b/a")`)) })
```

Sie sollten `x` aus Ihrer `solve`-Methode zurückgeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_return("x")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
--fcc-editable-region--
    def solve(self):
        pass
--fcc-editable-region--
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
print(lin_eq)

```
