---
id: 663b7fefd437bd984e091cbf
title: Schritt 29
challengeType: 20
dashedName: step-29
---

# --description--

Erstellen Sie als Nächstes eine neue Klasse namens `QuadraticEquation` und lassen Sie sie von `Equation` erben. Sie verwenden diese neue Klasse, um quadratische Gleichungen darzustellen, also Gleichungen zweiten Grades in der Form $ax^2 + bx + c = 0$. 

Definieren Sie innerhalb Ihrer neuen Klasse ein `degree`-Klassenattribut mit dem Wert `2`, der dem Grad einer quadratischen Gleichung entspricht. Definieren Sie außerdem die `solve`- und `analyze`-Methoden. Die Implementierung übernehmen Sie in den folgenden Schritten.

# --hints--

Sie sollten eine neue Klasse namens `QuadraticEquation` erstellen und diese von der Klasse `Equation` erben lassen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").inherits_from("Equation")`)) })
```

Sie sollten eine `solve`-Methode innerhalb der `QuadraticEquation`-Klasse definieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").has_function("solve")`)) })
```

Ihre `solve`-Methode sollte einen einzelnen Parameter `self` annehmen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("solve").has_args("self")`)) })
```

Sie sollten eine `analyze`-Methode innerhalb der `QuadraticEquation`-Klasse definieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").has_function("analyze")`)) })
```

Ihre `analyze`-Methode sollte einen einzelnen Parameter `self` annehmen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("analyze").has_args("self")`)) })
```

Sie sollten ein `degree`-Klassenattribut innerhalb der `QuadraticEquation`-Klasse definieren und ihm den Wert `2` zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_variable("degree").is_equivalent("degree = 2")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}
--fcc-editable-region--

--fcc-editable-region--
lin_eq = LinearEquation(2, 3)
print(lin_eq)

```
