---
id: 663b93aee129b3c4cc07d0db
title: Schritt 20
challengeType: 20
dashedName: step-20
---

# --description--

Direkt nach der `terms`-Liste erstellen Sie eine `for`-Schleife und verwenden die `.items()`-Methode, um die in dem `coefficients`-Attribut gespeicherten Schlüssel und Werte zu durchlaufen. Verwenden Sie `n` und `coefficient` als Schleifenvariablen.

Innerhalb der Schleife erstellen Sie eine `if`-Anweisung, die überprüft, ob der Koeffizient bei der aktuellen Iteration einen falsy-Wert hat, und überspringen in diesem Fall die Iteration. Das liegt daran, dass Sie Koeffizienten mit dem Wert Null nicht darstellen möchten.

# --hints--

Sie sollten eine `for`-Schleife erstellen, die `coefficients.items()` durchläuft.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_for_iter().is_equivalent("self.coefficients.items()")`)) })
```

Ihre `for`-Schleife sollte `n` und `coefficient` verwenden, um `coefficients.items()` zu durchlaufen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_for_vars().is_equivalent("n, coefficient")`)) })
```

Sie sollten eine `if`-Anweisung erstellen, um zu überprüfen, ob `coefficient` innerhalb Ihrer `for`-Schleife einen falsy-Wert hat.

```js
({ test: () => assert(runPython(`
if_cond = _Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[0].find_conditions()[0]
conditions = ["not coefficient", "coefficient == 0", "0 == coefficient"]
any(if_cond.is_equivalent(condition) for condition in conditions)
`)) })
```

Sie sollten das Schlüsselwort `continue` innerhalb Ihrer neuen `if`-Anweisung verwenden.

```js
({ test: () => assert(runPython(`
_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[0].find_bodies()[0].has_stmt("continue")
`)) })
```

Ihre `for`-Schleife sollte direkt nach der Deklaration von `terms` platziert werden.

```js
({ test: () => assert(runPython(`
loop = str(_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0])
_Node(_code).find_class("Equation").find_function("__str__").is_ordered("terms = []", loop, "equation_string = ' '.join(terms)", "return equation_string")
`)) })
```


# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    def __str__(self):
        terms = []
        
--fcc-editable-region--
        equation_string = ' '.join(terms)
        return equation_string        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
