---
id: 663b83a28943e6aa6275a514
title: Schritt 19
challengeType: 20
dashedName: step-19
---

# --description--

Definieren Sie noch innerhalb der `Equation`-Klasse eine `__str__`-Methode, um den Gleichungsobjekten, die Sie erstellen werden, eine passende String-Darstellung zu geben.

Deklarieren Sie vorerst innerhalb der `__str__`-Methode eine Variable `terms` und weisen Sie ihr eine leere Liste zu. Diese Variable verwenden Sie, um jeden Term (Koeffizient mal \\( x^n \\)) Ihrer Gleichung zu speichern.

Deklarieren Sie dann eine Variable `equation_string`, weisen Sie ihr das Ergebnis zu, das entsteht, wenn Sie die Elemente der `terms`-Liste mit einem Leerzeichen verbinden. Geben Sie schließlich `equation_string` zurück.

# --hints--

Sie sollten eine `__str__`-Methode innerhalb der `Equation`-Klasse definieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__str__")`)) })
```

Ihre `__str__`-Methode sollte einen Parameter `self` annehmen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_args("self")`)) })
```

Sie sollten innerhalb der `__str__`-Methode eine Variable `terms` deklarieren und ihr eine leere Liste zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("terms = []")`)) })
```

Sie sollten innerhalb der `equation_string`-Methode eine Variable `terms` deklarieren und ihr das Ergebnis zuweisen, das entsteht, wenn Sie die Elemente in `__str__` mit einem Leerzeichen verbinden.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("equation_string = ' '.join(terms)")`)) })
```

Sie sollten aus Ihrer `equation_string`-Methode `__str__` zurückgeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_return("equation_string")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    
--fcc-editable-region--    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
