---
id: 663c9f31306353460da54542
title: Schritt 23
challengeType: 20
dashedName: step-23
---

# --description--

Wie Sie sehen, fehlt das `+`-Zeichen in der Ausgabe. Das Nummernzeichen wird standardmäßig nur bei negativen Zahlen angezeigt. Um dieses Verhalten zu ändern, können Sie nach dem auszuwertenden Ausdruck innerhalb der geschweiften Klammern Ihres f-Strings einen Doppelpunkt schreiben und die Option `+` angeben. So können Sie das Vorzeichen sowohl für positive als auch für negative Zahlen anzeigen.

Ändern Sie den String in Ihren beiden bedingten Klauseln, indem Sie `:+` innerhalb der geschweiften Klammern nach `coefficient` hinzufügen.

# --hints--

Sie sollten den String ändern, der an die `terms`-Liste innerhalb Ihrer `if`-Anweisung angehängt wird, und zwar in `f'{coefficient:+}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[1].find_bodies()[0].is_equivalent("terms.append(f'{coefficient:+}')")`)) })
```

Sie sollten den String ändern, der in die `terms`-Liste innerhalb Ihrer `elif`-Klausel eingefügt wird, und zwar in `f'{coefficient:+}x'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[1].find_bodies()[1].is_equivalent("terms.append(f'{coefficient:+}x')")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient}')
            elif n == 1:
                terms.append(f'{coefficient}x')                
        equation_string = ' '.join(terms)        
        return equation_string
--fcc-editable-region--    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
