---
id: 664eef158d792a509e8d708a
title: Schritt 34
challengeType: 20
dashedName: step-34
---

# --description--

Die `sub`-Funktion aus dem `re`-Modul ermöglicht es Ihnen, Text innerhalb eines Strings basierend auf einem Regex-Muster zu ersetzen.

```py
verse = 'Always look on the bright side of life'
spam = re.sub('bright', 'spam', verse)
spam == 'Always look on the spam side of life' # True
```

Sie nimmt drei Argumente entgegen: das Regex-Muster zum Finden, die Ersetzung und den String, auf dem Sie die Ersetzung durchführen möchten.

Geben Sie aus Ihrer `__str__`-Methode einen `sub()`-Aufruf zurück, der den String `'1'`, einen leeren String und Ihren bestehenden `equation_string.strip('+')`-Aufruf als Argumente übergibt. Dadurch wird jeder `1` durch einen leeren String ersetzt. Das Ergebnis ist noch nicht verfeinert, und Sie werden im nächsten Schritt weiter am Regex-Muster arbeiten.

# --hints--

Sie sollten aus Ihrer `re.sub()`-Methode einen `__str__`-Aufruf zurückgeben. Übergeben Sie dem `'1'`-Aufruf den String `equation_string.strip('+')`, einen leeren String und Ihren bestehenden `re.sub()`-Aufruf als Argumente.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_return("re.sub('1', '', equation_string.strip('+'))")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
import re


class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')
            else:
                terms.append(f"{coefficient:+}x**{n}")
        equation_string = ' '.join(terms) + ' = 0'
--fcc-editable-region--
        return equation_string.strip('+')        
--fcc-editable-region--
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}

class QuadraticEquation(Equation):
    degree = 2

    def __init__(self, *args):
        super().__init__(*args)
        a, b, c = self.coefficients.values()
        self.delta = b**2 - 4 * a * c
   
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
quadr_eq = QuadraticEquation(11, -1, 1)
print(quadr_eq)

```
