---
id: 66759e32b88fb5459b1e0234
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Die `__init_subclass__`-Methode wird immer dann aufgerufen, wenn die Klasse, die sie definiert, unterklassenartig erweitert wird, und sie ermöglicht es, die Kindklassen anzupassen. Die Methode nimmt einen Parameter mit dem konventionellen Namen `cls` (steht für „class“) entgegen, der die neue Kindklasse darstellt.

Definieren Sie eine `__init_subclass__`-Methode in Ihrer `Equation`-Klasse und geben Sie ihr einen `cls`-Parameter.

# --hints--

Sie sollten eine `__init_subclass__`-Methode mit einem `cls`-Parameter in Ihrer `Equation`-Klasse definieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").has_args("cls")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    
--fcc-editable-region--
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
