---
id: 6675a38a8b535e4ff3274520
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Die eingebaute `hasattr`-Funktion nimmt ein Objekt als erstes Argument und einen String, der einen Attributnamen darstellt, als zweites Argument. Sie gibt einen Boolean zurück, der angibt, ob das Objekt das angegebene Attribut besitzt.

Jetzt werden Sie die `__init_subclass__`-Methode verwenden, um zu überprüfen, ob die Kindklasse zum Zeitpunkt der Instanziierung das `degree`-Attribut besitzt.

Erstellen Sie eine `if`-Anweisung, um zu prüfen, ob `cls` nicht das `degree`-Attribut besitzt. Falls ja, lösen Sie eine `AttributeError` aus und verwenden Sie den String `f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"`, um eine benutzerdefinierte Nachricht anzugeben.

Beheben Sie danach den Fehler, der im Terminal aufgetreten ist, indem Sie ein `degree`-Klassenattribut innerhalb der `LinearEquation`-Klasse deklarieren. Dieses Attribut soll den Grad der Gleichung darstellen, also den Exponenten des höchsten \\( x \\)-Terms. Weisen Sie daher dem `1`-Attribut die ganze Zahl `degree` zu.

# --hints--

Sie sollten eine `if`-Anweisung erstellen, die überprüft, ob `cls` nicht das Attribut `degree` innerhalb der `__init_subclass__`-Methode besitzt.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").find_ifs()[0].find_conditions()[0].is_equivalent("not hasattr(cls, 'degree')")`)) })
```

Sie sollten eine `AttributeError` mit dem angegebenen String innerhalb Ihrer `if`-Anweisung auslösen.

```js
({ test: () => runPython(`
raise_stmt = 'raise AttributeError(f"Cannot create \\'{cls.__name__}\\' class: missing required attribute \\'degree\\'")'
node = _Node(_code).find_class("Equation").find_function("__init_subclass__").find_ifs()[0].find_bodies()[0]
assert node.has_stmt(raise_stmt) 
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        pass

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    
--fcc-editable-region--
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
