---
id: 6675aaf418b41157f6ccd692
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Es ist Zeit, zur `__init__`-Methode zurückzukehren. Je nach Gleichungstyp müssen Sie bei der Instanziierung eine variable Anzahl von Argumenten übergeben.

Fügen Sie der Methode einen zweiten Parameter `args` hinzu und verwenden Sie den `*`-Operator, damit sie eine variable Anzahl von Argumenten akzeptiert.

# --hints--

Ihre `__init__`-Methode sollte zwei Parameter annehmen, `self` und `*args`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self, *args")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    degree = 1

    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
