---
id: 690ddb5e6ee94eb73ed172b0
title: 기부 양식 디버그하기
challengeType: 25
dashedName: debug-donation-form
# Purposefully removed the demo type because this is a debugging project.
---

# --description--

지역 자선단체가 기부 양식 웹사이트를 만들었지만, 수정해야 할 여러 문제가 있습니다. 양식은 접근성이 없고 HTML 구문 오류가 있습니다.

모든 오류를 고쳐서 양식이 제대로 작동하고 모든 사용자가 접근할 수 있도록 하세요. 아래 사용자 스토리의 항목을 완료하고 "Run the Tests"를 클릭해 모든 오류를 고쳤는지 확인하세요.

**사용자 스토리:**

1. `input` 요소는 빈 요소이므로 닫는 태그가 없어야 합니다. 양식에서 모든 `</input>` 닫는 태그를 제거하세요.
2. 각 입력 필드에 `label` 요소를 추가해 사용자가 각 필드가 무엇을 위한 것인지 알 수 있게 하세요. 레이블 텍스트는 현재 각 입력 옆에 있는 텍스트와 일치해야 합니다.
3. `Email Address:` 입력 타입은 `email` 대신 `text`여야 합니다.
4. 각 `label` 요소를 해당 `input` 요소와 `for` 속성을 `label`에, 일치하는 `id` 속성을 `input`에 사용해 연결해야 합니다.
5. 텍스트, 이메일, 숫자 입력 필드(체크박스나 제출 버튼 제외)에 `required` 속성을 추가해 사용자가 필수 정보를 입력하도록 하세요.

# --hints--

코드에 `</input>` 닫는 태그가 없어야 합니다.

```js
assert.notMatch(code, /<\/input>/);
```

양식에 정확히 다섯 개의 `input` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("input"), 5);
```

양식에 정확히 네 개의 `label` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

첫 번째 `label`에는 `Full Name:` 텍스트가 있어야 합니다.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[0]?.textContent, /Full\s*Name\s*:/i);
```

첫 번째 `input`에는 `required` 속성이 있어야 합니다.

```js
const firstInput = document.querySelector("input[name='name']");
assert.isTrue(firstInput?.hasAttribute("required"));
```

두 번째 `label`에는 `Email Address:` 텍스트가 있어야 합니다.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[1]?.textContent, /Email\s*Address\s*:/i);
```

두 번째 `input`에는 `required` 속성이 있어야 합니다.

```js
const secondInput = document.querySelector("input[name='email']");
assert.isTrue(secondInput?.hasAttribute("required"));
```

세 번째 `label`에는 `Donation Amount ($):` 텍스트가 있어야 합니다.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[2]?.textContent, /Donation\s*Amount\s*\(\$\)\s*:/i);
```

세 번째 `input`에는 `required` 속성이 있어야 합니다.

```js
const thirdInput = document.querySelector("input[name='amount']");
assert.isTrue(thirdInput?.hasAttribute("required"));
```

네 번째 `label`에는 `Subscribe` 텍스트가 있어야 합니다.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[3]?.textContent, /Subscribe/i);
```

모든 `input` 요소를 `label` 요소와 `for`, `id` 속성을 사용해 연결해야 합니다.

```js 
const els = document.querySelectorAll('input:not([type="submit"])'); 
assert.isNotEmpty(els); 
els.forEach(input => { 
  const label = document.querySelector(`label[for="${input.id}"]`); 
  assert.exists(label);  
  assert.equal(label.tagName, "LABEL"); 
})
``` 

이메일 주소용 `input`는 `email` 타입이어야 합니다.

```js
const emailInput = document.querySelector("input[name='email']");
assert.strictEqual(emailInput?.getAttribute("type"), "email");
```

체크박스 `input`에는 `required` 속성이 없어야 합니다.

```js
const checkboxInput = document.querySelector("input[name='newsletter']");
assert.isFalse(checkboxInput?.hasAttribute("required"));
```

제출 버튼 `button`에는 `required` 속성이 없어야 합니다.

```js
const submitButton = document.querySelector("input[type='submit']");
assert.isFalse(submitButton?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    
    Full Name:
    <input type="text" name="name"></input>

    Email Address:
    <input type="text" name="email">

    Donation Amount ($):
    <input type="number" name="amount"></input>

    <input type="checkbox" name="newsletter"></input>
    Subscribe

    <input type="submit" value="Send"></input>
  </form>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    <label for="full-name">Full Name:</label>
    <input id="full-name" type="text" name="name" required>

    <label for="email">Email Address:</label>
    <input id="email" type="email" name="email" required>

    <label for="amount">Donation Amount ($):</label>
    <input id="amount" type="number" name="amount" required>

    <label for="newsletter">Subscribe</label>
    <input id="newsletter" type="checkbox" name="newsletter">

    <input type="submit" value="Send">
  </form>
</body>
</html>
```
