---
id: 68d275dd800f404d22a07564
title: 深さ優先探索アルゴリズムの実装
challengeType: 27
dashedName: lab-depth-first-search
---

# --description--

このラボでは、グラフの探索アルゴリズムである<dfn>深さ優先探索</dfn>を実装します。

<dfn>幅優先探索</dfn>がソースノードからの辺の長さを段階的に探索するのに対し、<dfn>深さ優先探索</dfn>はまず一つの辺の経路をできるだけ深く進みます。

経路の端に到達すると、探索は未訪問の辺の経路がある最後のノードまで戻り、探索を続けます。

幅優先探索とは異なり、ノードを訪問するたびにすべての隣接ノードを訪問するわけではありません。代わりに、まず一つの隣接ノードを訪問し、その経路を辿って訪問可能なノードがなくなるまで進みます。

このアルゴリズムを実装するには、スタック（最後に追加された要素が最初に取り出される配列、<dfn>後入れ先出し</dfn>の原則に従う）を使うとよいでしょう。スタックは深さ優先探索で役立ちます。なぜなら、隣接ノードをスタックに追加するとき、最新に追加された隣接ノードを優先して訪問し、スタックから取り出したいからです。

このアルゴリズムの単純な出力は、あるノードから到達可能なノードのリストです。したがって、訪問したノードを記録しておく必要があります。

**目的**：以下のユーザーストーリーを満たし、すべてのテストを通過してラボを完了してください。

**ユーザーストーリー：**

1. `dfs`という名前の関数が必要です。
1. `dfs`関数は2つの引数を取る必要があります：
   - 無向の隣接行列。
   - `0`と`n - 1`の間のノードの数値ラベル。ここで`n`はグラフのノード総数です。
1. `dfs`関数は、渡されたノードから到達可能なすべてのノードのリストを出力する深さ優先探索アルゴリズムを実装する必要があります。

# --hints--

`dfs`という名前の関数が2つの引数を取る必要があります。

```js
({ test: () => runPython(`
  import inspect
  assert inspect.isfunction(dfs)
  sig = inspect.signature(dfs)
  assert len(sig.parameters) == 2
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)`は`1`、`2`、`3`、`0`を含むリストを返すべきです。

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)
  assert isinstance(result, list)
  assert len(result) == 4
  assert set(result) == {1, 2, 3, 0}
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)`は`1`、`2`、`3`、`0`を含むリストを返すべきです。

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)
  assert isinstance(result, list)
  assert len(result) == 4
  assert set(result) == {3, 2, 1, 0}
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)`は`[3]`を返すべきです。

```js
({ test: () => runPython(`
  assert dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3) == [3]
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)`は`3`と`2`を含むリストを返すべきです。

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)
  assert result == [3, 2] or result == [2, 3]
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)`は`0`と`1`を含むリストを返すべきです。

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)
  assert result == [0, 1] or result == [1, 0]
`) })
```

`dfs`関数は正しい結果を返すべきです。

```js
({ test: () => runPython(`
  import random

  def solve(g, s):
      v, q = set(), [s]
      while q:
          u = q.pop()
          if u not in v:
              v.add(u)
              q.extend([i for i, x in enumerate(g[u]) if x == 1 and i not in v])
      return v

  random.seed(0)

  for _ in range(10):
      n = 6
      mat = [[0]*n for _ in range(n)]
      for i in range(n):
          for j in range(i+1, n):
              if random.random() > 0.5:
                  mat[i][j] = mat[j][i] = 1

      start = random.randint(0, n-1)
      assert set(dfs(mat, start)) == solve(mat, start), "Random graph check failed"
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def dfs(graph, root):
    stack = []
    temp_v = None
    visited = []
    temp_v_neighbors = []
    stack.append(root)
    while stack:
        temp_v = stack.pop()
        if temp_v not in visited:
            visited.append(temp_v)
            temp_v_neighbors = graph[temp_v]
            for n, is_neighbor in enumerate(temp_v_neighbors):
                if is_neighbor == 1:
                    stack.append(n)
    return visited
```
