---
id: 642ddfdea4200e313f80a4b6
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

请记住，`document` 对象有一个 `.createElement()` 方法，它允许你动态创建新的 HTML 元素。

在 `createLabel` 函数中，声明一个 `label` 变量并为其分配一个新的 `div` 元素。

# --hints--

你应该在 `createLabel` 函数中声明一个 `label` 变量。

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*(?:const|let|var)\s+label/);
```

你的 `label` 变量应使用 `const` 声明。

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*const\s+label/);
```

你应该使用 `document` 对象的 `.createElement()` 方法。

```js
assert.match(code, /document\.createElement\(/);
```

你应该将字符串 `"div"` 传递给 `.createElement()` 方法。

```js
assert.match(code, /document\.createElement\(\s*('|"|`)div\1\s*\)/);
```

你应该将新的 `div` 元素分配给 `label`。

```js
assert.match(code, /const\s+label\s*=\s*document\.createElement\(\s*('|"|`)div\1\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {

  }
}
--fcc-editable-region--
```
