---
id: 642df9df4b5216350de7b0d2
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

你将需要一个函数来生成一系列数字。

声明一个空的 `range` 函数，该函数接受 `start` 和 `end` 参数。 使用 `Array()` 构造函数并隐式返回一个空数组。

# --hints--

你应该声明一个 `range` 变量。

```js
assert.match(code, /(?:let|var|const)\s+range/);
```

你应该使用 `const` 来声明你的 `range` 变量。

```js
assert.match(code, /const\s+range/);
```

你的 `range` 变量应该是一个函数。

```js
assert.isFunction(range);
```

你的 `range` 函数应该使用箭头语法。

```js
assert.match(code, /const\s+range\s*=\s*\(.*\)\s*=>/);
```

你的 `range` 函数应该首先接受 `start` 参数。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start/)
```

你的 `range` 函数应该第二接受 `end` 参数。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

你的 `range` 函数应该使用隐式返回。 请记住，这意味着你将不会使用花括号。

```js
assert.notMatch(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

你的 `range` 函数应该使用 `Array()` 构造函数。 原始构造函数不需要 `new` 关键字。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(/);
```

你不应将任何内容传递给 `Array()` 构造函数。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
