---
id: 642db8c409d9991d0b3b2f0d
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

你的项目从一个基本的 HTML 容器和一些相应的 CSS 开始。 你的首要任务是以编程方式生成电子表格的单元格。

全局 `window` 对象代表浏览器窗口（或选项卡）。 它有一个 `onload` 属性，允许你定义窗口加载整个页面（包括样式表和脚本）时的行为。

首先将 `window` 的 `onload` 属性设置为没有参数的箭头函数。 在函数中声明一个 `container` 变量，并为其赋值通过 `id` 为 `"container"` 获取元素的值。

# --hints--

你应该访问 `window` 对象的 `onload` 属性。

```js
assert.match(code, /window\.onload/);
```

你应该将 `onload` 属性设置为一个函数。

```js
assert.isFunction(window.onload);
```

你应该使用箭头语法。

```js
assert.match(code, /window\.onload\s*=\s*\(/);
```

你的 `onload` 函数不应接受任何参数。

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)/);
```

你应该在 `onload` 函数中声明一个 `container` 变量。

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*(?:let|var|const)\s+container/);
```

你的 `container` 变量应使用 `const` 声明。

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container/);
```

你应该使用 `document.getElementById()`

```js
assert.match(code, /document\.getElementById\(/);
```

你应该有一个 `id` 为 `"container"` 的元素。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

你应该将 `#container` 元素分配给 `container`。

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
