---
id: 642e02be7845f13b014cd2b0
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

现在你有了 `range` 函数，你也可以使用它来创建字母范围。

使用 `const` 和箭头语法声明 `charRange` 函数。 它应该采用 `start` 和 `end` 参数。 该函数应隐式返回以 `start` 和 `end` 作为参数调用 `range()` 的结果。

# --hints--

你应该声明一个 `charRange` 变量。

```js
assert.match(code, /(?:let|const|var)\s+charRange/);
```

你的 `charRange` 变量应使用 `const` 声明。

```js
assert.match(code, /const\s+charRange/);
```

你的 `charRange` 变量应该是一个函数。

```js
assert.isFunction(charRange);
```

你的 `charRange` 函数应该使用箭头语法。

```js
assert.match(code, /const\s+charRange\s*=\s*\(.*\)\s*=>/);
```

你的 `charRange` 函数应该将 `start` 作为第一个参数。

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start/);
```

你的 `charRange` 函数应该将 `end` 作为第二个参数。

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

你的 `charRange` 函数应该使用隐式返回。

```js
assert.notMatch(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

你的 `charRange` 函数应该调用你的 `range` 函数。

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(/);
```

你应该将 `start` 和 `end` 作为参数传递给 `range` 调用。

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\s*,\s*end\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
