---
id: 64345c560591891f64976f7a
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

现在你的辅助函数已完成，回到你的 `onload` 事件处理程序中，你应该声明一个 `letters` 变量。 将以字母 `"A"` 和 `"J"` 作为参数调用 `charRange()` 的结果赋值给它。

# --hints--

你应该声明一个 `letters` 变量。

```js
assert.match(code, /(?:let|const|var)\s+letters/);
```

你应该使用 `const` 来声明你的 `letters` 变量。

```js
assert.match(code, /const\s+letters/);
```

你应该为你的 `letters` 变量分配一个 `charRange()` 调用。

```js
assert.match(code, /const\s+letters\s*=\s*charRange\(/);
```

你应该将 `"A"` 作为第一个参数传递给 `charRange()` 调用。

```js
assert.match(code, /const\s+letters\s*=\s*charRange\(\s*('|"|`)A\1/);
```

你应该将 `"J"` 作为第二个参数传递给 `charRange()` 调用。

```js
assert.match(code, /const\s+letters\s*=\s*charRange\(\s*('|"|`)A\1\s*,\s*('|"|`)J\2\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
