---
id: 643498755d54c6279ba09078
title: 步骤 20
challengeType: 0
dashedName: step-20
---

# --description--

大多数电子表格程序都包含用于计算的内置函数。

声明一个 `sum` 函数，该函数接受一个 `nums` 参数，该参数是一个数字数组。 它应该返回对数组调用 `reduce` 对所有数字求和的结果。

# --hints--

你应该声明一个 `sum` 变量。

```js
assert.match(code, /(?:let|const|var)\s+sum/);
```

你应该使用 `const` 来声明你的 `sum` 变量。

```js
assert.match(code, /const\s+sum/);
```

你的 `sum` 变量应该是一个函数。

```js
assert.isFunction(sum);
```

你的 `sum` 函数应该使用箭头语法。

```js
assert.match(code, /const\s+sum\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

你的 `sum` 函数应该有一个 `nums` 参数。

```js
assert.match(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>/);
```

你的 `sum` 函数应该使用隐式返回。

```js
assert.notMatch(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*{/);
```

你的 `sum` 函数应该返回在 `nums` 上调用 `.reduce()` 的结果。

```js
assert.match(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*nums\.reduce\(/);
```

你的 `sum` 函数应该返回 `nums` 中所有数字的总和。

```js
const numbers = [1, 2, 3, 4, 5];
assert.equal(sum(numbers), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
