---
id: 643715013330824ecaa70442
title: 步骤 23
challengeType: 0
dashedName: step-23
---

# --description--

你的下一个函数将计算数字数组的中值。 首先声明一个接受 `nums` 参数的 `median` 箭头函数。

在函数中，声明一个 `sorted` 变量，并为其分配对 `nums` 数组的副本进行排序的值。

你应该使用 `slice()` 方法来创建数组的浅拷贝。

# --hints--

你应该声明一个 `median` 变量。

```js
assert.match(code, /(?:let|const|var)\s+median/);
```

你应该使用 `const` 来声明你的 `median` 变量。

```js
assert.match(code, /const\s+median/);
```

你的 `median` 变量应该是一个函数。

```js
assert.isFunction(median);
```

你的 `median` 函数应该使用箭头语法。

```js
assert.match(code, /const\s+median\s*=\s*\(?/);
```

你的 `median` 函数应该有一个 `nums` 参数。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)/);
```

你的 `median` 函数不应使用隐式返回。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{/);
```

你的 `median` 函数应该有一个 `sorted` 变量。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*(?:let|var|const)\s+sorted/);
```

你应该使用 `const` 来声明你的 `sorted` 变量。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted/);
```

你应该使用 `.slice()` 将 `nums` 数组的副本分配给 `sorted`。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)/);
```

你应该将 `.sort()` 方法链接到 `.slice()` 方法。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(/);
```

你应该将回调函数传递给 `sort` 方法，以便准确地按升序对数字进行排序。 为了清楚起见，使用隐式返回。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(\s*\(\s*a\s*,\s*b\s*\)\s*=>\s*a\s*-\s*b\s*\)\s*\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
