---
id: 641d9a19bff38d34d5a5edb8
title: Hatua ya 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utafundishwa jinsi ya kufanya kazi na Fetch API ili kupakua data kutoka chanzo cha nje na kuionyesha kwenye ukurasa uliojaa waandishi wa freeCodeCamp.

HTML na CSS zote za warsha hii zimekwisha tolewa kwako. Unaweza kuangalia mafaili hayo mawili ili kuzoea.

Unapokuwa tayari, anza kwa kupata upatikanaji wa vipengele vya `#author-container` na `#load-more-btn`. Kumbuka unaweza kutumia njia za `getElementById` au `querySelector` kwa hili.

Kisha weka thamani ya vipengele hivi kwenye vigezo vinavyoitwa `authorContainer` na `loadMoreBtn`, kwa mpangilio huo.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `authorContainer`.

```js
assert.isDefined(authorContainer);
```

Unapaswa kuweka thamani ya kipengele cha `#author-container` kwenye kigezo cha `authorContainer`.

```js
assert.strictEqual(authorContainer, document.querySelector('#author-container'));
```

Unapaswa kuwa na kigezo kinachoitwa `loadMoreBtn`.

```js
assert.isDefined(loadMoreBtn);
```

Unapaswa kuweka thamani ya kipengele cha `#load-more-btn` kwenye kigezo cha `loadMoreBtn`.

```js
assert.strictEqual(loadMoreBtn, document.querySelector('#load-more-btn'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
