---
id: 56533eb9ac21ba0edf2244b3
title: конвертувати Цельсій у Фаренгейт
challengeType: 1
forumTopicId: 16806
dashedName: convert-celsius-to-fahrenheit
---

# --description--

Формула для конвертації з Цельсія у Фаренгейт — це температура в Цельсіях, помножена на `9/5`, плюс `32`.

Вам задано змінну `celsius`, яка представляє температуру в Цельсіях. Використайте вже визначену змінну `fahrenheit` і призначте їй еквівалентну температуру у Фаренгейтах для заданої температури в Цельсіях. Використайте наведену вище формулу, щоб допомогти конвертувати температуру з Цельсія у Фаренгейт.

# --hints--

`convertCtoF(0)` має повертати число

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` має повертати значення `-22`

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` має повертати значення `14`

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` має повертати значення `32`

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` має повертати значення `68`

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` має повертати значення `86`

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js
function convertCtoF(celsius) {
  let fahrenheit;
  return fahrenheit;
}

convertCtoF(30);
```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
