---
id: a202eed8fc186c8434cb6d61
title: Реверс рядка
challengeType: 1
forumTopicId: 16043
dashedName: reverse-a-string
---

# --description--

Переверніть наданий рядок і поверніть перевернутий рядок.

Наприклад, `"hello"` має стати `"olleh"`.

# --hints--

`reverseString("hello")` має повертати рядок.

```js
assert.isString(reverseString('hello'));
```

`reverseString("hello")` має повертати рядок `olleh`.

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

`reverseString("Howdy")` має повертати рядок `ydwoH`.

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

`reverseString("Greetings from Earth")` має повертати рядок `htraE morf sgniteerG`.

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js
function reverseString(str) {
  return str;
}

reverseString('hello');
```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
