---
id: ab6137d4e35944e21037b769
title: Регістр заголовка для речення
challengeType: 1
forumTopicId: 16088
dashedName: title-case-a-sentence
---

# --description--

Поверніть наданий рядок, у якому перша літера кожного слова буде великою. Переконайтеся, що решта літер у слові — у нижньому регістрі.

Для цього завдання також потрібно робити великою літеру у сполучних словах, таких як `the` і `of`.

# --hints--

`titleCase("I'm a little tea pot")` має повертати рядок.

```js
assert.isString(titleCase("I'm a little tea pot"));
```

`titleCase("I'm a little tea pot")` має повертати рядок `I'm A Little Tea Pot`.

```js
assert.strictEqual(titleCase("I'm a little tea pot"), "I'm A Little Tea Pot");
```

`titleCase("sHoRt AnD sToUt")` має повертати рядок `Short And Stout`.

```js
assert.strictEqual(titleCase('sHoRt AnD sToUt'), 'Short And Stout');
```

`titleCase("HERE IS MY HANDLE HERE IS MY SPOUT")` має повертати рядок `Here Is My Handle Here Is My Spout`.

```js
assert.strictEqual(
  titleCase('HERE IS MY HANDLE HERE IS MY SPOUT'),
  'Here Is My Handle Here Is My Spout'
);
```

# --seed--

## --seed-contents--

```js
function titleCase(str) {
  return str;
}

titleCase("I'm a little tea pot");
```

# --solutions--

```js
function titleCase(str) {
  return str
    .split(' ')
    .map(word => word.charAt(0).toUpperCase() + word.substring(1).toLowerCase())
    .join(' ');
}

titleCase("I'm a little tea pot");
```
