---
id: 6690e10ebe2181212abc9652
title: 명함 디자인하기
challengeType: 25
dashedName: design-a-business-card
demoType: onClick
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 모든 다른 요소를 포함할 `div`가 `class` 속성에 `business-card` 값을 가져야 합니다.  
1. `.business-card` 요소 안에는 `img` 요소가 `class` 속성에 `profile-image` 값을 가져야 합니다. 원한다면 이미지 소스를 `https://cdn.freecodecamp.org/curriculum/labs/flower.jpg`로 설정할 수 있습니다. 의미 있는 설명을 가진 `alt`가 있어야 합니다.  
1. `.business-card` 요소 안에는 `img` 요소 다음에 각각 `p`, `class`, `full-name` 값을 가진 `designation` 속성을 가진 세 개의 `company` 요소가 있어야 합니다.  
1. 첫 번째 `p` 요소에는 이름이 들어가야 합니다.  
1. 두 번째 `p` 요소에는 직함이 들어가야 합니다.  
1. 세 번째 `p` 요소에는 회사 이름이 들어가야 합니다.  
1. 세 번째 `hr` 요소 아래에 `p` 요소가 있어야 합니다.  
1. `hr` 요소 다음에 두 개의 `p` 요소가 있어야 합니다. 첫 번째 `p` 요소에는 이메일 주소가 텍스트로 있어야 합니다. 두 번째 `p` 요소에는 전화번호가 텍스트로 있어야 합니다.  
1. 두 개의 `p` 요소 다음에 `a` 요소가 있어야 하며, 클래스는 `portfolio-link`이고 텍스트는 `Portfolio`이며 유효한 URL에 링크되어야 합니다.  
1. `hr` 요소 다음에 포트폴리오 링크를 포함하는 `a` 요소가 있어야 합니다.  
1. `div` 요소가 또 있어야 하며 `class` 속성에 `social-media` 값을 가져야 합니다. 이 요소 안에는 `h2` 요소가 있고 텍스트는 `Connect with me`입니다.  
1. `.social-media` 요소 안에는 세 개의 `a` 요소가 있어야 합니다. `a` 요소들은 각각 `Twitter`, `LinkedIn`, `GitHub` 텍스트를 가지며 유효한 웹사이트에 링크되어야 합니다.  
1. HTML 파일의 `styles.css` 섹션에 `link` 태그를 사용해 스타일 시트 `head`를 링크해야 합니다.  
1. 페이지 배경은 `rosybrown`이어야 합니다. 전체 폰트는 `Arial`이고 대체 수단은 `sans-serif`이어야 합니다.  
1. `.business-card` 선택자는 요소 너비를 `300px`로 설정하고 원하는 배경색을 가져야 합니다. 또한 모든 방향에 `padding` 속성으로 `20px` 값을 설정하고 위쪽 여백은 `margin` 속성으로 `100px` 값을 가져야 합니다. 텍스트는 가운데 정렬되고 폰트 크기는 `16px`이어야 합니다. 좌우 여백은 `auto`로 설정해야 합니다.  
1. `.profile-image` 선택자는 `max-width` 속성에 `100%` 값을 가져야 합니다.  
1. 모든 `p` 요소는 위아래 `margin` 속성에 `5px` 값을 가져야 합니다.  
1. 페이지의 모든 링크는 밑줄이 없어야 합니다.

# --hints--

하나의 주요 `div`가 `class` 속성에 `business-card` 값을 가져야 합니다.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'business-card');
```

`.business-card` 요소 안에는 `img` 요소가 `class` 속성에 `profile-image` 값을 가져야 합니다.

```js
assert.exists(document.querySelector('img.profile-image'));
```

이미지 소스는 유효한 이미지로 설정해야 합니다.

```js
assert.isAbove(document.querySelector('.profile-image')?.getAttribute('src')?.length, 0);
```

이미지의 `alt` 속성은 의미 있는 설명으로 설정해야 합니다.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

`.business-card` 요소 안에는 `img` 요소 다음에 `p` 요소가 `class` 속성에 `full-name` 값을 가져야 합니다.

```js
assert.exists(document.querySelector('div.business-card p.full-name'));
```

`.full-name` 요소 다음에 `p` 요소가 `class` 속성에 `designation` 값을 가져야 합니다.

```js
assert.exists(document.querySelector('div.business-card p.designation'));
```

`.designation` 요소 다음에 `p` 요소가 `class` 속성에 `company` 값을 가져야 합니다.

```js
assert.exists(document.querySelector('div.business-card p.company'));
```

클래스가 `p`인 첫 번째 `full-name` 요소에는 이름이 들어가야 합니다.

```js
assert.isNotEmpty(document.querySelector(".full-name")?.innerText.trim());
```

`p` 속성이 `class`인 두 번째 `designation` 요소에는 직함이 들어가야 합니다.

```js
assert.isNotEmpty(document.querySelector(".designation")?.innerText.trim());
```

`p` 속성이 `class`인 세 번째 `company` 요소에는 회사 이름이 들어가야 합니다.

```js
assert.isNotEmpty(document.querySelector(".company")?.innerText.trim());
```

`p` 속성이 `class`인 세 번째 `company` 요소 다음에 `hr` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('p.company + hr'));
```

첫 번째 `hr` 요소 다음에 이메일 주소가 텍스트인 `p` 요소가 있어야 합니다.

```js
assert.isNotEmpty(document.querySelectorAll('.business-card hr + p')[0].innerText.trim());
```

이메일 `p` 요소 다음에 전화번호가 텍스트인 `p` 요소가 있어야 합니다.

```js
assert.match(document.querySelectorAll(".business-card p")[4].textContent,/.+/);
```

전화번호 `p` 요소 다음에 클래스가 `a`인 `portfolio-link` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('a.portfolio-link'));
```

`.portfolio-link` 요소는 `Portfolio` 텍스트를 가져야 합니다.

```js
assert.equal(document.querySelector('a.portfolio-link').innerText.trim(), 'Portfolio');
```

`.portfolio-link` 요소는 유효한 `href`를 가져야 합니다.

```js
assert.isAbove(document.querySelector('.portfolio-link').getAttribute('href').length, 0);
```

포트폴리오를 포함하는 `a` 요소 다음에 구분선 역할을 하는 `hr` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('a + hr'));
```

두 번째 `hr` 요소 다음에 `div` 요소가 `class` 속성에 `social-media` 값을 가져야 합니다.

```js
assert.exists(document.querySelector('div.business-card div.social-media'));
```

`.social-media` 요소 안에는 `h2` 요소가 있고 텍스트는 `Connect with me`입니다.

```js
assert.equal(document.querySelector('.social-media h2')?.textContent.trim(), 'Connect with me');
```

`.social-media` 요소 안에는 세 개의 `a` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('.social-media a'), 3);
```

첫 번째 `a` 요소는 `Twitter` 텍스트를 가져야 합니다.

```js
assert.equal(document.querySelectorAll(".social-media a")[0]?.textContent.trim(), 'Twitter');
```

첫 번째 `href` 요소의 `a`는 유효한 링크를 가리켜야 합니다.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[0]?.getAttribute('href')?.length, 0);
```

두 번째 `a` 요소는 `LinkedIn` 텍스트를 가져야 합니다.

```js

assert.equal(document.querySelectorAll(".social-media a")[1]?.textContent.trim(), 'LinkedIn');
```

두 번째 `href` 요소의 `a`는 유효한 링크를 가리켜야 합니다.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[1]?.getAttribute('href')?.length, 0);
```

세 번째 `a` 요소는 `GitHub` 텍스트를 가져야 합니다.

```js
assert.equal(document.querySelectorAll(".social-media a")[2]?.textContent.trim(), 'GitHub');
```

세 번째 `href` 요소의 `a`는 유효한 링크를 가리켜야 합니다.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[2]?.getAttribute('href')?.length, 0);
```

`link` 요소는 `head` 요소 안에 있어야 합니다.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

`link` 요소는 `rel` 속성에 `stylesheet` 값을 가져야 합니다.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

`link` 요소는 `href` 속성에 `styles.css` 값을 가져야 합니다.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');
```

페이지 배경색은 `body` 선택자에서 `rosybrown`로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background-color'), 'rosybrown');
```

페이지 폰트는 `Arial` 요소에서 `body`로 설정하고 대체 수단은 `sans-serif`로 설정해야 합니다.

```js
const el = new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('font-family');
const condition1= el === 'Arial, sans-serif';
const condition2= el === '"Arial", sans-serif';
assert.isTrue(condition1 || condition2);
```

`.business-card` 선택자는 `width` 속성에 `300px` 값을 가져야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('width'), '300px');
```

`.business-card` 선택자는 유효한 색상 값을 가진 `background-color` 속성을 가져야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.business-card')?.backgroundColor);
```

`.business-card` 선택자는 `padding` 속성에 `20px` 값을 가져야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('padding'), '20px');
```

`.business-card` 선택자는 위쪽 여백이 `100px`이어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-top'), '100px');
```

`.business-card` 요소 안의 텍스트는 가운데 정렬되어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('text-align'), 'center');
```

`.business-card` 요소 안의 텍스트 폰트 크기는 `16px`이어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('font-size'), '16px');
```

`.business-card` 요소의 좌우 여백은 `auto`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-left'), 'auto');
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-right'), 'auto');
```

`.profile-image` 선택자는 `max-width` 속성에 `100%` 값을 가져야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image')?.getPropVal('max-width'), '100%');
```

`p` 요소들은 위아래 `margin` 속성에 `5px` 값을 가져야 합니다.

```js
const marginTop = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-top');
const marginBottom = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-bottom');

assert.equal(marginTop, '5px');
assert.equal(marginBottom, '5px');
```

페이지의 링크는 밑줄이 없어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('a')?.getPropVal('text-decoration'), 'none');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Business Card</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Digital Business Card</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/flower.jpg" alt="A vibrant flower in full bloom" class="profile-image">
        <p class="full-name">Your Name</p>
        <p class="designation">Software Developer</p>
        <p class="company" >@freeCodeCamp</p>
        <hr>
        <p>Email: your-email@example.com</p>
        <p>Phone: (123) 456-7890</p>
        <a class="portfolio-link" href="http://example.com">Portfolio</a>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    max-width: 100%;
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    margin: 20px 0;
}
```
