---
id: 69a03ddce82b0626f329b001
title: 15단계
challengeType: 1
dashedName: step-15
---

# --description--

이제 `while` 반복문을 만들어 `current`이 엄격하게 `null`가 아니고 `current.element`이 요소와 엄격하게 같지 않은지 확인하세요.

즉, 리스트의 끝에 도달하거나 제거할 요소를 찾을 때까지 반복문이 리스트를 계속 탐색해야 합니다.

# --hints--

조건 `while`를 가진 `current !== null && current.element !== element` 반복문을 만들어야 합니다.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /while\s+\(current\s*!==\s*null\s*&&\s*current\.element\s*!==\s*element\)\s*\{\s*/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    --fcc-editable-region--
        
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
