---
id: 69a0102e761a245bb1cbd8e7
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

이 시점에서 연결 리스트를 더 잘 이해하기 위해 테스트해 보세요.

`myList` 변수를 만들고 `initList()`로 설정하세요. 그 아래에 `isEmpty()`를 매개변수로 사용하여 `myList`를 호출한 결과를 콘솔에 기록하세요. 그런 다음 `add`와 원하는 숫자를 매개변수로 하여 `myList`를 호출하고, `myList`를 콘솔에 기록하세요. 그 아래에는 다시 `isEmpty()`를 매개변수로 사용하여 `myList`를 호출한 결과를 콘솔에 기록하세요.

# --hints--

`myList` 변수를 만들고 `initList()`로 설정하세요.

```js
assert.deepEqual(myList, initList());
```

`isEmpty(myList)`를 호출한 결과를 콘솔에 기록하세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(isEmpty\(myList\)\);?/)
```

`add`와 원하는 숫자를 매개변수로 하여 `myList`를 호출하세요.

```js
assert.match(__helpers.removeJSComments(code), /add\s*\(\s*myList\s*,\s*\d+\s*\)/);
```

`myList` 변수를 콘솔에 기록하세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?/)
```

`isEmpty(myList)`를 다시 콘솔에 기록하세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?\s*console\.log\(isEmpty\(myList\)\);?/)
```


# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
