---
id: 69a014bd9b1730e94136279f
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

리스트가 비어 있지 않으면, 리스트의 끝을 찾아 새 요소를 추가하고 그것이 `null`를 가리키도록 해야 합니다. 이를 위해 `head`부터 리스트를 반복해야 합니다.

이를 시작하려면 `else` 절을 열고, 그 안에서 `let`를 사용해 `current` 변수를 만들고 `list.head`로 설정하세요.

# --hints--

`else` 절이 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{/);
```

`current` 절 안에 `list.head`로 설정된 `else` 변수를 만들어야 합니다.

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{\s*let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    --fcc-editable-region--
    } 
    --fcc-editable-region--

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
