---
id: 69a045fd236d0d973a59f103
title: 19단계
challengeType: 1
dashedName: step-19
---

# --description--

만약 `previous`가 `null`라면, 이는 제거할 요소가 리스트의 맨 앞에 있다는 뜻이므로, 첫 번째 요소를 제거하려면 `head`가 다음 노드를 가리키도록 업데이트해야 합니다.

이를 위해 `else` 절을 추가하세요. 그 안에서 `list.head`를 `current.next`로 설정하세요.

# --hints--

`else` 문장에 `if` 절을 추가해야 합니다.

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*/);
```

`list.head` 절 안에서 `current.next`를 `else`로 설정해야 합니다.

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*list\.head\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }

    if(previous !== null) {
        previous.next = current.next;
    --fcc-editable-region--
    } 
        
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
