---
id: 69a013a86c97dac3784b2636
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

이제 콘솔에 `{ head: { element: 42, next: null }, length: 1 }` 같은 내용이 표시될 것입니다.

하지만 추가된 노드가 리스트의 첫 번째 요소, 즉 헤드인 경우를 이렇게 처리하지는 않습니다.

그렇게 하려면, `isEmpty`를 인자로 하여 `list` 함수를 호출해 리스트가 비어 있는지 if 문으로 확인하세요. 만약 참이라면, `list.head = node` 줄을 `if` 문 안으로 옮기세요.

# --hints--

조건 `if`를 가진 `isEmpty(list)` 문이 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*/);
```

`list.head = node` 줄을 `if` 문 안으로 옮기세요.

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*list\.head\s*=\s*node;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

--fcc-editable-region--
    list.head = node;
--fcc-editable-region--
    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
