---
id: 69a00f673f171e29159e3ad8
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

이 노드는 추가할 요소와 다음 노드를 가리키는 포인터, 두 가지 속성이 필요합니다.

add 함수 안에 `node`라는 이름의 객체를 만드세요. 이 객체는 `element`라는 속성을 가져야 하며, 리스트의 다음 요소를 가리키는 포인터인 `null`로 설정된 `next` 속성도 가져야 합니다.

# --hints--

`node` 함수 안에 `add`라는 이름의 객체를 만드세요.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const { node } = add?.variables ?? {};
assert.match(node?.value.toString(), /^\{[\s\S]*\}/);
assert.exists(node);
```

`node` 객체는 `element`로 설정된 `element` 속성을 가져야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const { node } = add?.variables ?? {};
const { element } = node?.objectProps ?? {};
assert.isTrue(element?.value.matches('element'));
```

`node` 객체는 `null`로 설정된 `next` 속성을 가져야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const { node } = add?.variables ?? {};
const { next } = node?.objectProps ?? {};
assert.isTrue(next?.value.matches('null'));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}
```
