---
id: 69a0179a85b00b6127e0af45
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

`while` 반복문 안에서 `current`을 `current.next`로 업데이트하여 반복문이 리스트를 한 노드씩 앞으로 이동하게 하고, `next` 속성이 `null`인 노드, 즉 리스트의 끝에 도달하면 정지하게 하세요.

# --hints--

`current` 반복문 안에서 `current.next` 변수의 값을 `while`로 업데이트해야 합니다.

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            --fcc-editable-region--

            --fcc-editable-region--
        }
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
