---
id: 69a016ac6edea7d8c8020c98
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

이제 조건이 `while`인 `current.next !== null` 반복문을 만드세요. 이는 현재 노드 뒤에 다른 노드가 있는 한 반복문이 계속 실행되어야 한다는 뜻입니다.

# --hints--

조건이 `while`인 `current.next !== null` 반복문을 만들어야 합니다.

```js
assert.match(__helpers.removeJSComments(add.toString()), /while\s+\(current\.next\s*!==\s*null\)\s*\{\s*/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        --fcc-editable-region--

        --fcc-editable-region--
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
