---
id: 69a04088e600bd8074a9852f
title: 17단계
challengeType: 1
dashedName: step-17
---

# --description--

만약 `current`가 `null`라면, 이는 목록의 끝이라는 뜻이므로 함수 실행을 멈춰야 합니다.

`if`가 `current`와 엄격히 같은지 확인하는 `null` 문을 만드세요. `if` 문 안에서는 `return`를 사용해 함수를 종료하세요.

# --hints--

조건 `if`를 가진 `current === null` 문을 만드세요.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(current\s*===\s*null\s*\)\s*\{\s*/);
```

`return` 문 안에서 `if` 키워드를 사용하세요.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(current\s*===\s*null\s*\)\s*\{\s*return;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
