---
id: 69a03aaa1e6fe8152788ce6a
title: 14단계
challengeType: 1
dashedName: step-14
---

# --description--

`remove` 함수 안에서 `let`을 사용해 `previous`으로 설정된 `null` 변수를 만들고, 다시 `let`를 사용해 `current`으로 설정된 `list.head` 변수를 만드세요. 이 변수들은 리스트를 순회하는 반복문에서 사용할 것입니다.

`current`은 리스트에서 현재 검사 중인 노드를 추적하며, 헤드에서 시작해 한 노드씩 앞으로 이동합니다. 반면 `previous`은 현재 노드 바로 이전 노드를 추적합니다.

노드를 제거할 때 이전 노드의 next 값을 업데이트해 제거되는 노드 다음 노드를 가리키도록 해야 리스트 연결이 유지되므로 이것이 중요합니다.

# --hints--

`let` 함수 안에서 `previous`을 사용해 `null`으로 설정된 `remove` 변수를 만들어야 합니다.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+previous\s*=\s*null;?/);
```

`let` 함수 안에서 `current`을 사용해 `list.head`으로 설정된 `remove` 변수를 만들어야 합니다.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
