---
id: 69a03f70b50b4c4d9a0f0113
title: 16단계
challengeType: 1
dashedName: step-16
---

# --description--

루프 안에서 `previous`를 `current`로, `current`를 `current.next`로 다시 할당하세요. 이는 `current`가 다음 노드를 확인하고, `previous`는 항상 한 노드 뒤에 머무른다는 뜻입니다.

# --hints--

`previous` 변수를 `current`로 다시 할당해야 합니다.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /previous\s*=\s*current;?/);
```

`current` 변수를 `current.next`로 다시 할당해야 합니다.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /current\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) { 
        --fcc-editable-region--
                
        --fcc-editable-region--
    }
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
