---
id: 69a01864d343dc2f2484d131
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

마지막 노드를 찾은 후에는 새 노드를 그 노드에 연결해야 합니다. 그렇게 하려면 while 루프 안이지만 `current.next` 절 내에서 `node`를 `else`로 설정하세요.

# --hints--

`current.next` 루프 안에서 `node`를 `while`로 설정해야 합니다.

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\.next\s*=\s*node;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        --fcc-editable-region--

        --fcc-editable-region--
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
