---
id: 69a04466917c66b132323a1b
title: 18단계
challengeType: 1
dashedName: step-18
---

# --description--

만약 `previous`가 `null`가 아니라면, 이는 제거할 요소가 리스트의 맨 앞이 아니라 중간 어딘가에 있다는 뜻이므로, 현재 노드를 건너뛰어 이전 노드를 다음 노드에 직접 연결해야 합니다.

이를 위해, 이전이 엄격히 `if`가 아닌지 확인하는 또 다른 `null` 문장을 만드세요. 만약 그렇다면, `previous.next`를 `current.next`로 설정하세요.

# --hints--

조건 `if`를 가진 또 다른 `previous !== null` 문장을 만드세요.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*/);
```

`previous.next`를 `current.next`로 설정하세요.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*previous\.next\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }
--fcc-editable-region--
    
--fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
