---
id: 69a012c6224fac85832247dd
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

콘솔에 나타난 `true`, `{ head: null, length: 0 }`, 그리고 `true` 결과를 확인하세요. 이는 리스트 `myList`에 아무것도 할당하지 않았고, `length` 속성을 증가시키지 않았기 때문입니다.

`node`에 `list.head`를 할당하고, 그 아래에서 `length`의 `list` 속성을 증가시키세요.

# --hints--

`node`에 `list.head`를 할당해야 합니다.

```js
assert.match(__helpers.removeJSComments(add.toString()), /list\.head\s*=\s*node;?/);
```

`length`의 `list` 속성을 증가시켜야 합니다.

```js
assert.match(__helpers.removeJSComments(add.toString()), /list\.length\+\+;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    --fcc-editable-region--

    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
