---
id: 69a00cd4ec513345d40e95ec
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

다음으로 작업할 것은 목록이 비어 있는지 확인할 수 있는 함수입니다.

이를 위해 `isEmpty`라는 이름의 함수를 만들고 `list` 매개변수를 받으세요. 함수 내부에서 `length`의 `list` 속성이 `0`와 엄격히 같은지 확인하는 값을 반환하세요.

# --hints--

`isEmpty` 함수를 만들어야 합니다.

```js
assert.isFunction(isEmpty)
```

`isEmpty` 함수는 `list` 매개변수를 가져야 합니다.

```js
const isEmptyParams = __helpers.getFunctionParams(isEmpty.toString());
assert.deepEqual(isEmptyParams[0].name, "list");
```

`isEmpty` 함수는 `list.length === 0`를 반환해야 합니다.

```js
const emptyList = { head: null, length: 0 };
const nonEmptyList = { head: {}, length: 1 };

assert.isTrue(isEmpty(emptyList));
assert.isFalse(isEmpty(nonEmptyList));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
