---
id: 69a02ee82bbeaee00e44c67d
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

목록에 요소를 추가하려면 이 정도면 충분합니다. 어떻게 보이는지 확인하려면 원하는 숫자 두 개를 `myList` 목록에 더 추가하세요.

# --hints--

원하는 숫자 두 개를 `myList` 목록에 더 추가하세요.

```js
const matches = __helpers.removeJSComments(code).match(/add\s*\(\s*myList\s*,\s*\d+\s*\);?/g) || [];

assert.isAtLeast(matches.length, 3);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
--fcc-editable-region--

--fcc-editable-region--
console.log(myList)
console.log(isEmpty(myList))
```
