---
id: 69a00eaa43458e618331873c
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

이제 리스트에 노드를 추가하는 함수를 작성할 차례입니다.

`add` 함수를 `list`와 `element` 매개변수로 만드세요.

# --hints--

`add` 함수를 만들어야 합니다.

```js
assert.isFunction(add)
```

`add` 함수는 `list`와 `element` 매개변수를 가져야 합니다.

```js
const addParams = __helpers.getFunctionParams(add.toString());
const addParamsNames = addParams.map(param => param.name)
assert.include(addParamsNames, "list");
assert.include(addParamsNames, "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
