---
id: 69a00a051b9a18694d00c16b
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 연결 리스트를 구현할 것입니다.

먼저, `initList`라는 이름의 함수를 만드세요. 그 안에서 `head` 속성이 `null`로 설정된 객체와 `length` 속성이 `0`인 객체를 반환하세요.

이 함수는 연결 리스트를 초기화할 때 호출할 함수입니다. `head`는 리스트의 첫 번째 요소를 나타내며, 처음에는 리스트에 요소가 없기 때문에 `null`입니다. `length`은 리스트에 있는 노드(요소)의 수를 추적하며, 노드를 추가하거나 제거할 때 이 값을 증가시키거나 감소시킬 것입니다.

# --hints--

`initList` 함수를 만들어야 합니다.

```js
assert.isFunction(initList)
```

`initList` 함수는 객체를 반환해야 합니다.

```js
assert.isObject(initList());
```

`initList` 함수는 `head`이 `null`로 설정된 객체를 반환해야 합니다.

```js
assert.propertyVal(initList(), "head", null);
```

`initList` 함수는 `length`가 `0`로 설정된 객체를 반환해야 합니다.

```js
assert.propertyVal(initList(), "length", 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
