---
id: 69a030fbe487f185f55b8627
title: 13단계
challengeType: 1
dashedName: step-13
---

# --description--

이제 리스트에서 노드를 제거하는 함수를 작업할 것입니다.

`remove` 함수를 `list`와 `element` 매개변수로 만드세요.

# --hints--

`remove` 함수를 만들어야 합니다.

```js
assert.isFunction(remove)
```

`remove` 함수는 `list`와 `element` 매개변수를 가져야 합니다.

```js
const removeParams = __helpers.getFunctionParams(remove.toString());
const removeParamsNames = removeParams.map(param => param.name)
assert.include(removeParamsNames, "list");
assert.include(removeParamsNames, "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

--fcc-editable-region--

--fcc-editable-region--

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
