---
id: 65f569725359e10d345bc52a
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

Побудуйте програму, яка розраховує та малює траєкторію снаряда якій дано кут, швидкість та висоту кидка.

Для початку імпортуйте `math`, ви часто використовуватимете його у цьому проєкті, оскільки він містить такі корисні методи, як `math.radians`, `math.cos`, `math.sin` та інші.

Також створіть ці змінні, щоб вони мали значення гравітаційного прискорення та деякі спеціальні символи, які будуть корисними пізніше (використайте для них копіювання та вставлення).

```py
GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
```

# --hints--

Використайте `import math`, щоб імпортувати модуль `math`.

```js
({test: () => assert(runPython(`_Node(_code).has_import('import math')`))})
```

Ви повинні мати змінну `GRAVITATIONAL_ACCELERATION` зі значенням `9.81`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("GRAVITATIONAL_ACCELERATION"), "missing variable"
assert _Node(_code).find_variable("GRAVITATIONAL_ACCELERATION").is_equivalent('GRAVITATIONAL_ACCELERATION = 9.81'), "variable has wrong value"
`)})
```

Ви повинні мати змінну `PROJECTILE` зі значенням `"∙"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("PROJECTILE"), "missing variable"
assert _Node(_code).find_variable("PROJECTILE").is_equivalent('PROJECTILE = "∙"'), "variable has wrong value"
`)})
```

Ви повинні мати змінну `x_axis_tick` зі значенням `"T"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("x_axis_tick"), "missing variable"
assert _Node(_code).find_variable("x_axis_tick").is_equivalent('x_axis_tick = "T"'), "variable has wrong value"
`)})
```

Ви повинні мати змінну `y_axis_tick` зі значенням `"⊣"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("y_axis_tick"), "missing variable"
assert _Node(_code).find_variable("y_axis_tick").is_equivalent('y_axis_tick = "⊣"'), "variable has wrong value"
`)})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
