---
id: 65f569725359e10d345bc52a
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

Ви збираєтеся створити програму, яка обчислює та малює траєкторію снаряда, враховуючи кут, швидкість і висоту кидка.

Почніть з імпорту `math`, ви часто будете його використовувати в цьому проєкті, оскільки він має корисні методи, такі як `math.radians`, `math.cos`, `math.sin` та інші.

Також створіть ці змінні, щоб мати значення прискорення вільного падіння та деякі спеціальні символи, які знадобляться пізніше (використовуйте копіювання та вставлення для них).

```py
GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
```

# --hints--

Використовуйте `import math` для імпорту модуля `math`.

```js
({test: () => assert(runPython(`_Node(_code).has_import('import math')`))})
```

У вас має бути змінна `GRAVITATIONAL_ACCELERATION` зі значенням `9.81`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("GRAVITATIONAL_ACCELERATION"), "missing variable"
assert _Node(_code).find_variable("GRAVITATIONAL_ACCELERATION").is_equivalent('GRAVITATIONAL_ACCELERATION = 9.81'), "variable has wrong value"
`)})
```

У вас має бути змінна `PROJECTILE` зі значенням `"∙"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("PROJECTILE"), "missing variable"
assert _Node(_code).find_variable("PROJECTILE").is_equivalent('PROJECTILE = "∙"'), "variable has wrong value"
`)})
```

У вас має бути змінна `x_axis_tick` зі значенням `"T"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("x_axis_tick"), "missing variable"
assert _Node(_code).find_variable("x_axis_tick").is_equivalent('x_axis_tick = "T"'), "variable has wrong value"
`)})
```

У вас має бути змінна `y_axis_tick` зі значенням `"⊣"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("y_axis_tick"), "missing variable"
assert _Node(_code).find_variable("y_axis_tick").is_equivalent('y_axis_tick = "⊣"'), "variable has wrong value"
`)})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
