---
id: 65fd4b3bedee044b4b957d33
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Створіть клас `Projectile` з методом `__init__`, який ініціалізує клас за допомогою трьох аргументів: початкової швидкості, початкової висоти та початкового кута кидка снаряда, у цьому порядку.

Всередині методу `__init__` призначте ці аргументи трьом приватним атрибутам: `__speed`, `__height` та `__angle`. Кут, що задається, буде в градусах; однак його слід зберігати внутрішньо в радіанах. Для цього використайте функцію `math.radians()` для конвертації кута з градусів у радіани під час призначення його `__angle`.

Використання двох підкреслень перед ім’ям атрибута викликає манглінг імен у Python. Це означає, що атрибути не доступні безпосередньо ззовні класу за їхніми заданими іменами, і до них потрібно звертатися через заманглені імена, як `ball._Projectile__height`, якщо потрібно звернутися зовні, що вказує на те, що вони призначені лише для внутрішнього використання.

# --hints--

Ви повинні використати `class Projectile:` для створення класу.

```js
({test: () => assert(runPython(`_Node(_code).has_class("Projectile")`))})

```

У класі `__init__` має бути метод `Projectile`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_body().has_function("__init__")`))})
```

Новий екземпляр класу `Projectile` повинен мати правильно ініціалізовані значення трьох атрибутів.

```js
({test: () => assert(runPython(`
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
p = Projectile(20, 21, 22)
p._Projectile__height == 21 and p._Projectile__speed == 20 and p._Projectile__angle == math.radians(22)
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--

--fcc-editable-region--
```
