---
id: 65fd4bd84561a14d3e720061
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Змінна класу `__slots__` має особливе використання в класах Python. Оголошення `__slots__` та призначення до неї послідовності рядків обмежує створення атрибутів лише тими, що входять до цієї послідовності. Також, це запобігає створенню спеціального атрибута `__dict__` та забезпечує більш ефективному доступу до атрибутів.

Використайте змінну `__slots__` в межах класу, щоб визначити, які атрибути має клас: призначте до `__slots__` кортеж, що містить 3 рядки, кожен з яких дорівнює одному з імен атрибутів, визначених у `__init__`.

# --hints--

Призначте `__slots__` кортеж, який містить `'__height'`, `'__speed'` та `'__angle'`.

```js
({test: () => assert(runPython(`
from itertools import permutations
slots = _Node(_code).find_class("Projectile").find_body().find_variable("__slots__")
perms = permutations(("__height", "__speed", "__angle"))
any(slots.is_equivalent(f'__slots__ = {perm}') for perm in perms)
`))})

```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--
class Projectile:

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
--fcc-editable-region--
```
