---
id: 65fd4b3bedee044b4b957d33
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Створіть клас `Projectile` за допомогою методу `__init__`, який ініціалізує клас використовуючи три аргументи: початкова швидкість, початкова висота та початковий кут кидка снаряда, у цьому порядку.

Усередині методу `__init__` призначте ці аргументи трьом приватним атрибутам: `__speed`, `__height` та `__angle`. Наданий кут буде в градусах; однак його непогано було б зберігати внутрішньо в радіанах. Щоб досягти цього, скористайтеся функцією `math.radians()`, щоб перетворити кут із градусів у радіани під час призначення його до `__angle`.

Використання двох знаків підкреслення перед назвою атрибута запускає спотворення імені в Python. Це означає, що атрибути не доступні безпосередньо ззовні класу використовуючи їх імена, і до них потрібно звертатися за допомогою спотворених імен, таких як `ball._Projectile__height`, якщо це необхідно зовні, вказуючи, що вони призначені лише для внутрішнього використання.

# --hints--

Використайте `class Projectile:`, щоб створити клас.

```js
({test: () => assert(runPython(`_Node(_code).has_class("Projectile")`))})

```

Ви повинні мати метод `__init__` в межах класу `Projectile`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_body().has_function("__init__")`))})
```

Новий екземпляр класу `Projectile` повинен мати значення трьох атрибутів правильно ініціалізованими.

```js
({test: () => assert(runPython(`
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
p = Projectile(20, 21, 22)
p._Projectile__height == 21 and p._Projectile__speed == 20 and p._Projectile__angle == math.radians(22)
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--

--fcc-editable-region--
```
