---
id: 6616ec0bf9ee8548ce6b0f08
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Тепер час працювати над обчисленням координат траєкторії! Створіть метод класу `Projectile` з назвою `__calculate_y_coordinate`, який, крім `self`, має приймати параметр `x`.

\\[ y = y_0 + x \tan\theta - \frac{g x^2}{2 v_0^2 \cos^2\theta} \\]

Вищенаведена формула обчислює вертикальну позицію $y$ для будь-якої горизонтальної позиції $x$, маючи початковий кут $\theta$, швидкість $v_0$ та висоту $y_0$.

Вам потрібно буде використати `math.tan()` та `math.cos()`, і пам’ятайте, що `x ** y` — це спосіб записати $x^y$, а значення $g$ знаходиться у змінній `GRAVITATIONAL_ACCELERATION`.

Реалізуйте метод так, щоб він повертав координату $y$.

Якщо хочете протестувати метод, ззовні класу можна викликати `ball._Projectile__calculate_y_coordinate()` з числом як аргумент після того, як `ball` буде оголошено.

# --hints--

Ваш клас `Projectile` має мати метод з назвою `__calculate_y_coordinate`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').has_function('__calculate_y_coordinate')`))}})
```

Метод `__calculate_y_coordinate` має приймати `self, x` як аргументи.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').find_function('__calculate_y_coordinate').has_args('self, x')`))}})
```

Метод має повертати правильне значення. Для `Projectile(12, 12, 12)`, маючи 4 $x$ координати 0, 1, 2 і 3, координата y має бути приблизно 12, 12.18, 12.28, 12.32.

```js
({test: () => {
    assert(runPython(`
    a = Projectile(12, 12, 12)
    all([round(a._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(0, 12.0), (1, 12.18), (2, 12.28), (3, 12.32)]])
    `))
}})
```

Метод має повертати правильне значення. Для `Projectile(45, 12, 22)`, маючи $x$ координати 6, 9, 10 і 12, координата y має бути приблизно 14.32, 15.41, 15.76, 16.44.

```js
({test: () => {
    assert(runPython(`
    b = Projectile(45, 12, 22)
    all([round(b._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(6, 14.32), (9, 15.41), (10, 15.76), (12, 16.44)]])
    `))
}})
```

Метод `__calculate_y_coordinate` не повинен округлювати результат.

```js
({test: () => assert.isFalse(runPython(`Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5) == round(Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5), 2)`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    --fcc-editable-region--
    
    --fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
   
```
