---
id: 662f9f65c50bd21d05ffdee9
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Визначте метод з іменем `calculate_all_coordinates`, який обчислює координати для всіх значень \\(x\\) від `0` до зсуву, округленого вгору (не включно), а потім повертає їх у вигляді списку кортежів `(x, y)`.

Потім викличте метод `calculate_all_coordinates` на екземплярі `ball` і призначте вихідні дані новій змінній з іменем `coordinates`.

Ви можете використати `math.ceil()` для округлення числа вгору до найменшого цілого числа, більшого або рівного цьому числу.

# --hints--

Ви повинні визначити метод `calculate_all_coordinates` з одним аргументом `self`.

```js
({test: () => assert(runPython(`
c = _Node(_code).find_class('Projectile')
c.has_function('calculate_all_coordinates') and c.find_function('calculate_all_coordinates').has_args('self')
`))})
```

Ви повинні викликати метод `calculate_all_coordinates` на `ball` і призначити результат змінній з іменем `coordinates`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('coordinates'), "coordinates variable is missing"

        assert _Node(_code).find_variable('coordinates').is_equivalent('coordinates = ball.calculate_all_coordinates()'), "wrong value assigned to coordinates variable"
    `)
})
```

Метод `calculate_all_coordinates` має повертати список, що містить усі координати `(x, y)`.

```js
({test: () => assert(runPython(`
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
ball = Projectile(12, 13, 14)
expected = [(x, ball._Projectile__calculate_y_coordinate(x)) for x in range(math.ceil(ball._Projectile__calculate_displacement()))]
actual = ball.calculate_all_coordinates()
expected == actual
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x**2 / (
                2 * self.__speed**2 * math.cos(self.__angle)**2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate

    --fcc-editable-region--
    
        
        
ball = Projectile(10, 3, 45)
print(ball)
   --fcc-editable-region--
```
