---
id: a26cbbe9ad8655a977e1ceb5
title: 最長単語検索アプリを作成する
challengeType: 26
dashedName: build-a-longest-word-finder-app
---

# --description--

このラボでは、与えられた文の中で最も長い単語の長さを返す関数を作成します。

例えば、文 `"The quick brown fox jumped over the lazy dog"` では、最も長い単語は `"jumped"` で、その長さは6です。

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー:**

1. `findLongestWordLength` という名前の関数を作成し、文字列を引数として受け取る必要があります。
2. その関数は、文字列の中で最も長い単語の長さを返す必要があります。

# --hints--

`findLongestWordLength` という名前の関数を作成してください。

```js
assert.isFunction(findLongestWordLength);
```

`findLongestWordLength` は1つのパラメータを持つ必要があります。

```js
assert.lengthOf(findLongestWordLength, 1);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")` は数値を返す必要があります。

```js
assert.isNumber(
  findLongestWordLength('The quick brown fox jumped over the lazy dog')
);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")` は `6` を返す必要があります。

```js
assert.strictEqual(
  findLongestWordLength('The quick brown fox jumped over the lazy dog'),
  6
);
```

`findLongestWordLength("May the force be with you")` は `5` を返す必要があります。

```js
assert.strictEqual(findLongestWordLength('May the force be with you'), 5);
```

`findLongestWordLength("Google do a barrel roll")` は `6` を返す必要があります。

```js
assert.strictEqual(findLongestWordLength('Google do a barrel roll'), 6);
```

`findLongestWordLength("Googling do a barrel roll")` は `8` を返す必要があります。

```js
assert.strictEqual(findLongestWordLength('Googling do a barrel roll'), 8);
```

`findLongestWordLength("What is the average airspeed velocity of an unladen swallow")` は `8` を返す必要があります。

```js
assert.strictEqual(
  findLongestWordLength(
    'What is the average airspeed velocity of an unladen swallow'
  ),
  8
);
```

`findLongestWordLength("What if we try a super-long word such as otorhinolaryngology")` は `19` を返す必要があります。

```js
assert.strictEqual(
  findLongestWordLength(
    'What if we try a super-long word such as otorhinolaryngology'
  ),
  19
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findLongestWordLength(str) {
  return str.split(' ').sort((a, b) => b.length - a.length)[0].length;
}

findLongestWordLength('The quick brown fox jumped over the lazy dog');
```
