---
id: 694bf6eb530e19a9c48c59f6
title: Schritt 17
challengeType: 20
dashedName: step-17
---

# --description--

Sie können von jedem Teil des Strings einen Ausschnitt nehmen, nicht nur vom Anfang. Das ist in vielen Fällen hilfreich.

Erstellen Sie eine Variable `year_code` und weisen Sie ihr den Ausschnitt von `employee_code` vom Index `4` bis `8` zu. Dadurch wird `2026` extrahiert.

Erstellen Sie dann eine Variable `initials` und weisen Sie ihr den Ausschnitt von `employee_code` vom Index `9` bis `11` zu. Dadurch wird `JD` extrahiert.

Geben Sie schließlich beide Variablen im Terminal aus.

# --hints--

Sie sollten eine Variable mit dem Namen `year_code` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("year_code")`)) })
```

Sie sollten `employee_code` vom Index `4` bis `8` ausschneiden und das Ergebnis Ihrer Variable `year_code` zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("year_code").is_equivalent("year_code = employee_code[4:8]")`)) })
```

Sie sollten eine Variable mit dem Namen `initials` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("initials")`)) })
```

Sie sollten `employee_code` vom Index `9` bis `11` ausschneiden und das Ergebnis Ihrer Variable `initials` zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("initials").is_equivalent("initials = employee_code[9:11]")`)) })
```

Sie sollten `year_code` im Terminal ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(year_code)")`)) })
```

Sie sollten `initials` im Terminal ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(initials)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department) 
--fcc-editable-region--  

--fcc-editable-region--  
```
