---
id: 694bf6ea530e19a9c48c59f0
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Wenn Sie zwei Strings wie `'John' + 'Doe'` aneinanderhängen, ist das Ergebnis `'JohnDoe'` ohne Leerzeichen. Um das zu beheben, müssen Sie einen String mit einem Leerzeichen (`' '`) dazwischen anfügen.

Aktualisieren Sie Ihre Variable `full_name`, sodass sie die Variable `first_name`, ein Leerzeichen (`' '`), und die Variable `last_name` in dieser Reihenfolge mit dem `+`-Operator verknüpft.

# --hints--

Sie sollten ein Leerzeichen (`' '`) zwischen `first_name` und `last_name` hinzufügen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + ' ' + last_name")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--
full_name = first_name + last_name
--fcc-editable-region--
print(full_name)
```
