---
id: 694bf6eb530e19a9c48c59f5
title: Schritt 16
challengeType: 20
dashedName: step-16
---

# --description--

Wenn Sie mit Strings arbeiten, müssen Sie oft einen bestimmten Teil eines Strings extrahieren. Das nennt man **Slicing**.

Die Syntax ist `string[start:stop]`, wobei:

* `start` der Index ist, an dem das Slice beginnt (**inklusive**).
* `stop` der Index ist, an dem das Slice endet (**exklusiv**).

Zum Beispiel, wenn `text = 'Python'`, dann gibt `text[0:2]` `'Py'`.

Definieren Sie `employee_code` als `'DEV-2026-JD-001'`. Erstellen Sie danach eine Variable `department` und weisen Sie ihr das Slice von `employee_code` von Index `0` bis `3` zu. Geben Sie dann `department` im Terminal aus.

# --hints--
Sie sollten eine Variable namens `employee_code` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_code")`)) })
```

Sie sollten den String `DEV-2026-JD-001` Ihrer Variable `employee_code` zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_code").is_equivalent("employee_code = 'DEV-2026-JD-001'")`)) })
```

Sie sollten eine Variable namens `department` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("department")`)) })
```

Sie sollten die ersten drei Zeichen von `employee_code` ausschneiden und sie Ihrer Variable `department` zuweisen.

```js
({ test: () => runPython(`
dep = _Node(_code).find_variable("department")
assert dep.is_equivalent("department = employee_code[0:3]") or dep.is_equivalent("department = employee_code[:3]")
`) })
```

Sie sollten `department` im Terminal ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(department)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)

--fcc-editable-region--

--fcc-editable-region--
```
