---
id: 694bf6eb530e19a9c48c59f7
title: Schritt 18
challengeType: 20
dashedName: step-18
---

# --description--

Sie können auch negative Zahlen verwenden, um von hinten an einem String zu schneiden. Zum Beispiel bezieht sich `-1` auf das letzte Zeichen, `-2` auf das vorletzte Zeichen und so weiter.

Um die letzten drei Zeichen eines Strings zu erhalten, können Sie `string[-3:]` verwenden. Beachten Sie, dass der Stopp-Parameter nach dem Doppelpunkt weggelassen wird. Das bedeutet, dass das Schneiden bis zur Grenze des Strings erfolgen soll.

Erstellen Sie eine Variable namens `last_three`. Verwenden Sie negative Indizes, um die letzten drei Zeichen von `employee_code` (die die ID-Nummer darstellen) zu extrahieren. Geben Sie schließlich `last_three` im Terminal aus.

Damit ist der Workshop zur Mitarbeiterausweis-Karte abgeschlossen.

# --hints--

Sie sollten eine Variable namens `last_three` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("last_three")`)) })
```

Sie sollten die letzten drei Zeichen von `employee_code` mit negativer Indizierung ausschneiden und sie Ihrer `last_three`-Variablen zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_three").is_equivalent("last_three = employee_code[-3:]")`)) })
```

Sie sollten `last_three` im Terminal ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_three)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
print(year_code)
initials = employee_code[9:11]
print(initials)
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
initials = employee_code[9:11]
print(year_code)
print(initials)
last_three = employee_code[-3:]
print(last_three)
```
