---
id: 694bf6eb530e19a9c48c59f2
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Vervollständigen Sie nun den Satz, indem Sie die `employee_info`-Zuweisung aktualisieren, sodass am Ende auch der String ` years old` angehängt wird. Denken Sie daran, am Anfang Ihres Strings ein Leerzeichen einzufügen.

Geben Sie abschließend `employee_info` aus.

# --hints--

Sie sollten `' years old'` an das Ende der Variablen `employee_info` anhängen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age) + ' years old'")`)) })
```

Sie sollten `employee_info` ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_info)")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + str(employee_age)

--fcc-editable-region-- 
```
