---
id: 694bf3af05d10fd89dab05cc
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

In Python können Sie Strings mit dem `+`-Operator kombinieren. Das nennt man Konkatenation. Hier ist ein Beispiel:

```python
greeting = 'Hello' + 'World'
print(greeting) # Output: HelloWorld
```

Erstellen Sie eine Variable `full_name`, indem Sie `first_name` und `last_name` verketten. Geben Sie dann `full_name` aus.

# --hints--

Sie sollten eine Variable mit dem Namen `full_name` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("full_name")`)) })
```

Sie sollten Ihrer `first_name`-Variable einen String zuweisen, der durch das Verketten von `last_name` und `full_name` gebildet wird.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + last_name")`)) })
```

Sie sollten die Variable `full_name` ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(full_name)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--

--fcc-editable-region--
```
