---
id: 694bf6eb530e19a9c48c59f4
title: Schritt 13
challengeType: 20
dashedName: step-13
---

# --description--

Viele Strings mit `+` zu verketten und Zahlen mit `str()` umzuwandeln, kann unübersichtlich und schwer lesbar werden.

Python 3.6 hat **f-Strings** eingeführt, um dieses Problem zu lösen. Indem Sie den Buchstaben `f` vor das öffnende Anführungszeichen setzen, können Sie Variablen und Ausdrücke in Ersetzungsfeldern darstellen, die durch geschweifte Klammern `{}` gekennzeichnet sind. Zum Beispiel:

```python
name = 'John'
print(f'Hello {name}') # Output: Hello John
```

Erstellen Sie eine Variable `employee_card` und weisen Sie ihr einen f-String zu, der `Employee:` gefolgt von einem Leerzeichen und dem Wert der Variablen `full_name` anzeigt.

# --hints--

Sie sollten eine Variable mit dem Namen `employee_card` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_card")`)) })
```

Ihre Variable `employee_card` sollte den Wert `f'Employee: {full_name}'` haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name}'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

--fcc-editable-region--
```
