---
id: 694bf6eb530e19a9c48c59f1
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Ihre Adresse scheint jetzt unvollständig zu sein. Sie möchten auch die Wohnungsnummer hinzufügen, in der der Mitarbeitende wohnt, also sollten Sie die Variable ändern.

Wenn Sie Inhalt an das Ende einer bestehenden String-Variable anhängen möchten, können Sie den **erweiterten Zuweisungs**operator `+=` verwenden. Das ist kürzer als `var = var + 'new text'` zu schreiben. Zum Beispiel:

```py
greeting = 'Hello'
greeting += ' World'
print(greeting) # Hello World
```

Denken Sie daran, dass Strings unveränderlich sind, daher ändert diese Operation den ursprünglichen String nicht. Stattdessen wird ein neuer String erstellt und der Variable neu zugewiesen.

Verwenden Sie den `+=`-Operator, um den String `, Apartment 4B` zu Ihrer `address`-Variable hinzuzufügen.

# --hints--

Sie sollten den `+=`-Operator verwenden, um den String `, Apartment 4B` zur `address`-Variable hinzuzufügen.

```js
({ test: () => assert(runPython(`_Node(_code).find_aug_variable("address").is_equivalent("address += ', Apartment 4B'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
address = '123 Main Street'  
--fcc-editable-region--  

--fcc-editable-region-- 
print(address)  
```
