---
id: 694bf6eb530e19a9c48c59f3
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Jetzt werden Sie die `str()`-Funktion noch einmal verwenden. Genau wie bei age müssen Sie jede numerische Variable in einen String umwandeln, bevor Sie sie mit anderem Text verketten.

Erstellen Sie eine Variable namens `experience_years` und weisen Sie ihr die Ganzzahl `5` zu.

Erstellen Sie dann eine Variable `experience_info`. Weisen Sie ihr einen String zu, der durch die Verkettung von `'Experience: '`, der in einen String umgewandelten `experience_years`-Variable und `' years'` gebildet wird. Geben Sie das Ergebnis im Terminal aus.

# --hints--

Sie sollten eine Variable namens `experience_years` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_years")`)) })
```

Sie sollten die Ganzzahl `5` Ihrer `experience_years`-Variable zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_years").is_equivalent("experience_years = 5")`)) })
```

Sie sollten eine Variable namens `experience_info` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_info")`)) })
```

Sie sollten Ihrer `'Experience: '`-Variable einen String zuweisen, der durch die Verkettung von `str(experience_years)`, `' years'` und `experience_info` gebildet wird.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_info").is_equivalent("experience_info = 'Experience: ' + str(experience_years) + ' years'")`)) })
```

Sie sollten `experience_info` ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(experience_info)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
--fcc-editable-region--

--fcc-editable-region--
```
