---
id: 695c9c7c61d446f1451e0ae7
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Erstellen Sie als Nächstes eine Variable `address`, um die Adresse des Mitarbeiters zu speichern. Weisen Sie ihr den String `123 Main Street` zu, und geben Sie schließlich `address` aus.

# --hints--

Sie sollten eine Variable mit dem Namen `address` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("address")`)) })
```

Sie sollten der Variablen `123 Main Street` den String `address` zuweisen. Denken Sie daran, den Text entweder in einfache oder doppelte Anführungszeichen zu setzen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("address").is_equivalent("address = '123 Main Street'")`)) })
```

Sie sollten die Variable `address` ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(address)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
--fcc-editable-region--

--fcc-editable-region--
```

