---
id: 695c9f21a8b1d53cf1b97841
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Erstellen Sie nun eine Variable mit dem Namen `employee_age` und weisen Sie ihr die Ganzzahl `28` zu. 


# --hints--

Sie sollten eine Variable mit dem Namen `employee_age` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_age")`)) })
```

Sie sollten die Ganzzahl `28` Ihrer Variable `employee_age` zuweisen. Umgeben Sie den Text nicht mit Anführungszeichen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_age").is_equivalent("employee_age = 28")`)) })
```


# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
--fcc-editable-region--

--fcc-editable-region--
```
