---
id: 695c96f82af4fb78e59fa9ed
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

Strings sind Folgen von Zeichen, die verwendet werden, um Textdaten zu speichern. Wie Sie sich vielleicht aus vorherigen Lektionen erinnern, können Sie einen String erstellen, indem Sie Text entweder in einfache (`'`) oder doppelte (`"`) Anführungszeichen einschließen. Zum Beispiel:

```python
greeting = 'Hello'
print(greeting) # Output: Hello
```

Erstellen Sie eine Variable `first_name`, die den String `John` speichert, und eine Variable `last_name`, die den String `Doe` speichert. Geben Sie dann `first_name` und `last_name` aus.

# --hints--

Sie sollten eine Variable mit dem Namen `first_name` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("first_name")`)) })
```

Sie sollten der Variablen `John` den String `first_name` zuweisen. Denken Sie daran, den Text entweder in einfache oder doppelte Anführungszeichen einzuschließen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("first_name").is_equivalent("first_name = 'John'")`)) })
```

Sie sollten eine Variable mit dem Namen `last_name` haben.

```js  
({ test: () => assert(runPython(`_Node(_code).has_variable("last_name")`)) })
```

Sie sollten der Variablen `Doe` den String `last_name` zuweisen. Denken Sie daran, den Text entweder in einfache oder doppelte Anführungszeichen einzuschließen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_name").is_equivalent("last_name = 'Doe'")`)) })
```

Sie sollten Ihre Variable `first_name` im Terminal ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(first_name)")`)) })
```

Sie sollten Ihre Variable `last_name` im Terminal ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_name)")`)) })
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
