---
id: 695df6480de23758af5a29cd
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Jetzt möchten Sie einen String erstellen, der das Alter des Mitarbeiters anzeigt.

Beginnen Sie damit, eine Variable `employee_info` zu erstellen und ihr das Ergebnis der Verkettung zuzuweisen:

- der Variable `full_name`.
- eines Strings, der aus den Zeichen `is` besteht, jeweils mit einem Leerzeichen davor und danach.

# --hints--

Sie sollten eine Variable namens `employee_info` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_info")`)) })
```

Sie sollten `full_name + ' is '` Ihrer Variable `employee_info` zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is '")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--

--fcc-editable-region--
```
