---
id: 695ea1e55c0e619e597fc5c9
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Wie Sie sehen, hat Python einen `TypeError: can only concatenate str (not "int") to str` ausgelöst. Das passiert, weil Python es nicht erlaubt, Text (Strings) und Zahlen (ganze Zahlen) direkt zu verketten.

Um das zu beheben, müssen Sie die Zahl zuerst mit der `str()`-Funktion in einen String umwandeln, die die String-Version eines Objekts zurückgibt:

```py
my_num = str(42)
print(type(my_num)) # <class 'str'>
```

Aktualisieren Sie Ihre `employee_info`-Zuweisung, um `employee_age` mit `str(employee_age)` in einen String umzuwandeln.

# --hints--

Sie sollten die Zuweisung Ihrer `employee_info`-Variablen ändern, indem Sie `employee_age` mit `str(employee_age)` in einen String umwandeln.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + employee_age
--fcc-editable-region--
```
