---
id: 695fd3dde03273875e5d6709
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Aktualisieren Sie die `employee_info`-Zuweisung, sodass `employee_age` am Ende angehängt wird.

Sobald Sie das getan haben, sehen Sie ein `TypeError` im Terminal. Im nächsten Schritt werden Sie daran arbeiten, es zu beheben.

# --hints--

Sie sollten die `employee_info`-Zuweisung aktualisieren, um `employee_age` am Ende anzuhängen.

```js
const regex = /^employee_info\s*=\s*full_name\s*\+\s*(["']) is \1\s*\+\s*employee_age\s*$/m;  
assert.match(__helpers.python.removeComments(code), regex);  
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is '
--fcc-editable-region--
```
