---
id: 69727ea5c9606bedfe36d3d2
title: Schritt 15
challengeType: 20
dashedName: step-15
---

# --description--

Jetzt ist es Zeit, die letzten Details zur Karte hinzuzufügen.

Erstellen Sie eine Variable namens `position` mit dem Wert des Strings `Data Analyst` und eine Variable namens `salary` mit dem Wert der Ganzzahl `75000`.

Aktualisieren Sie dann Ihre `employee_card`-f-String, um die Position und das Gehalt einzuschließen. Sie sollte genau folgendes Format haben: `Employee: [full_name] | Age: [employee_age] | Position: [position] | Salary: $[salary]`. Ersetzen Sie die Platzhalter durch die entsprechenden Variablen.

Geben Sie abschließend `employee_card` aus, um das Ergebnis zu sehen.

# --hints--

Sie sollten eine Variable namens `position` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("position")`)) })
```

Sie sollten den String `Data Analyst` Ihrer Variable `position` zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("position").is_equivalent("position = 'Data Analyst'")`)) })
```

Sie sollten eine Variable namens `salary` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("salary")`)) })
```

Sie sollten die Ganzzahl `75000` Ihrer Variable `salary` zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

Sie sollten `employee_card` aktualisieren, um den f-String `f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'` zu verwenden.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: \${salary}'")`)) })
```

Sie sollten `employee_card` im Terminal ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_card)")`))})
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

employee_card = f'Employee: {full_name} | Age: {employee_age}'
--fcc-editable-region--
```
