---
id: 6980d3ecfe800d8a54130f30
title: Schritt 14
challengeType: 20
dashedName: step-14
---

# --description--

Derzeit zeigt `employee_card` nur den Namen des Mitarbeiters an. Jetzt werden Sie weitere Informationen hinzufügen.

Aktualisieren Sie die `employee_card`-Zuweisung, um das Alter des Mitarbeiters einzuschließen. Der finale String sollte so aussehen: `Employee: [name] | Age: [age]` mit `[name]`, das durch den Namen des Mitarbeiters ersetzt wird, und `[age]`, das durch das Alter des Mitarbeiters ersetzt wird.

# --hints--

Sie sollten `employee_card` aktualisieren, um den f-String `f'Employee: {full_name} | Age: {employee_age}'` zu verwenden.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age}'")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--
employee_card = f'Employee: {full_name}'
--fcc-editable-region--
```
