---
id: 66eaddd04a9e533fba689001
title: Eine Blog-Post-Karte gestalten
challengeType: 25
dashedName: lab-blog-post-card
demoType: onClick
---

# --description--

In diesem Lab üben Sie, wie Sie Hintergründe und Rahmen gestalten, indem Sie eine Blog-Post-Karte erstellen.

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bestehen Sie alle Tests, um das Lab abzuschließen.

**User Stories:**

1. Sie sollten ein `div`-Element mit einer Klasse `blog-post-card` haben, das alle Elemente Ihrer Karte enthält.
2. Innerhalb des `.blog-post-card`-Elements sollten Sie ein Bild-Element mit einem gültigen `alt`-Attribut und Text sowie einer Klasse `post-img` haben. Sie können `https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg` für das `src`-Attribut des Bildes verwenden.
3. Sie sollten ein `div` mit einer Klasse `post-content` innerhalb des `.blog-post-card`-Elements haben mit Folgendem:
    - Ein `h2`-Element mit einer Klasse `post-title` und Text für den Titel Ihres Blog-Posts.
    - Ein `p`-Element mit einer Klasse `post-excerpt` und Text zur Zusammenfassung Ihres Blog-Posts.
    - Ein `a`-Element mit einer Klasse `read-more` mit dem Text `Read More`.
4. Sie sollten die Hintergrundfarbe des `body`-Elements auf einen anderen Wert als Weiß setzen.
5. Sie sollten die folgenden Stile auf das `.blog-post-card`-Element anwenden:
    - Einen weißen Hintergrund.
    - Abgerundete Ecken.
    - Eine Breite Ihrer Wahl.
    - Die von Ihnen gewählte Textausrichtung.
6. Das `.post-img`-Element sollte so gestylt sein, dass das Bild die gesamte Breite der Karte ausfüllt und einen unteren Rahmen hat.
7. Das `.post-content`-Element sollte so gestylt sein, dass sich Innenabstand innerhalb der Karte befindet.
8. Die `.post-title`- und `.post-excerpt`-Elemente sollten eine andere Textfarbe als die Standardfarbe und Außenabstände auf allen Seiten haben.
9. Das `.read-more`-Element sollte wie ein Button gestylt sein und folgende Eigenschaften haben:
    - Eine andere Textfarbe als die Standardfarbe.
    - Eine Hintergrundfarbe.
    - Außenabstände auf allen Seiten.
    - Innenabstände auf allen Seiten.
    - Abgerundete Ecken.
    - Eine `display`-Eigenschaft, die auf `inline-block` gesetzt ist.
    - Einen Hover-Effekt, der die Hintergrundfarbe ändert.

**Hinweis:** Vergewissern Sie sich, dass Sie Ihr Stylesheet in Ihrem HTML verlinken und Ihr CSS anwenden.

# --hints--

Sie sollten ein `div` mit einer Klasse `blog-post-card` haben.

```js
assert.exists(document.querySelector("div.blog-post-card"));
```

Sie sollten ein `img`-Element mit einem `class` von `post-img` haben. Stellen Sie sicher, dass Ihr Bild ein `alt`-Attribut mit Text und ein `src`-Attribut mit einem Wert hat.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const img = card.querySelector("img.post-img");
assert.exists(img);
assert.isString(img.alt);
assert.isNotEmpty(img.alt);
assert.isTrue(img.hasAttribute("src"));
assert.isNotEmpty(img.getAttribute("src"));
```

Sie sollten ein `div` mit der Klasse `post-content` haben.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);
const content = card.querySelector("div.post-content");
assert.exists(content);
```

Sie sollten ein `h2`-Element mit einer Klasse `post-title` haben. Stellen Sie sicher, dass es einen Textinhalt für den Titel Ihres Blog-Posts hat.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postTitle = content.querySelector("h2.post-title");
assert.exists(postTitle);
assert.isString(postTitle.textContent);
assert.isNotEmpty(postTitle.textContent);
```

Sie sollten ein `p` mit einer Klasse `post-excerpt` haben. Stellen Sie sicher, dass es einen Textinhalt zur Zusammenfassung Ihres Blog-Posts hat.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postExcerpt = content.querySelector("p.post-excerpt");
assert.exists(postExcerpt);
```

Sie sollten ein `a`-Element mit einem `class` von `read-more` haben.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const readMore = content.querySelector("a.read-more");
assert.exists(readMore);
```

Ihr `.read-more`-Element sollte den Text `Read More` haben.

```js
const el = document.querySelector("div.post-content a.read-more");
assert.exists(el);
assert.strictEqual(el.textContent.trim(), "Read More");
```

Ihr `.blog-post-card`-Element sollte eine `border-radius`-Eigenschaft mit einem Wert haben (sollte nicht `0` oder ein negativer Wert sein).

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const borderRadius = getComputedStyle(card).borderRadius;
assert.isAbove(parseInt(borderRadius), 0);
```

Ihr `.blog-post-card`-Element sollte einen weißen Hintergrund haben.

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const backgroundColor = getComputedStyle(card).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.oneOf(backgroundColor.toLowerCase(), whiteColors);
```

Ihr `body`-Element sollte eine andere Hintergrundfarbe als Weiß haben.

```js
const bodyBg = getComputedStyle(document.body).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.isFalse(whiteColors.includes(bodyBg.toLowerCase()));
```

Sie sollten `.blog-post-card` anvisieren und seine `width`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("width"));
```

Sie sollten `.blog-post-card` anvisieren und seine `text-align`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("text-align"));
```

Sie sollten `.post-content` anvisieren und seine `padding`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".post-content")?.getPropertyValue("padding"));
```

Ihr `.read-more`-Element sollte einen Hover-Effekt haben.

```js
const link = document.querySelector('.read-more');
assert.exists(link);

const d = new __helpers.CSSHelp(document).getStyle('.read-more:hover');
assert.isNotEmpty(d.backgroundColor);
```

Sie sollten `.read-more` anvisieren und seine `color`-Eigenschaft setzen.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreTextColor = getComputedStyle(readMore).color;
assert.notStrictEqual(readMoreTextColor, 'rgba(0, 0, 0, 0)');
```

Sie sollten `.read-more` anvisieren und seine `background-color`-Eigenschaft setzen.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBackgroundColor = getComputedStyle(readMore).backgroundColor;
assert.notStrictEqual(readMoreBackgroundColor, 'rgba(0, 0, 0, 0)');
```

Sie sollten `.read-more` anvisieren und seine `margin`-Eigenschaft setzen.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreMargin = getComputedStyle(readMore).marginTop;
assert.notStrictEqual(readMoreMargin, '0px');
```

Sie sollten `.read-more` anvisieren und seine `display`-Eigenschaft setzen.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreDisplay = getComputedStyle(readMore).display;
assert.strictEqual(readMoreDisplay, 'inline-block');
```

Sie sollten `.read-more` anvisieren und seine `border-radius`-Eigenschaft setzen.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBorderRadius = getComputedStyle(readMore).borderRadius;
assert.notStrictEqual(readMoreBorderRadius, '0px');
```

Sie sollten `.read-more` anvisieren und seine `padding`-Eigenschaft setzen.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMorePadding = getComputedStyle(readMore).padding;
assert.notStrictEqual(readMorePadding, '0px');
```

Ihr `.post-img`-Element sollte die `width` der Karte ausfüllen.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgWidth = getComputedStyle(postImg).width;
const cardWidth = getComputedStyle(postImg.parentElement).width;
assert.closeTo(Number(imgWidth.replace(/px/ig, "")), Number(cardWidth.replace(/px/ig, "")), 15);
```

Ihr `.post-img`-Element sollte einen `border-bottom`-Wert haben.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgBorderBottom = getComputedStyle(postImg).borderBottomWidth;
assert.notStrictEqual(imgBorderBottom, '0px');
```

Ihre `.post-title`- und `.post-excerpt`-Elemente sollten Außenabstände und nicht standardmäßige Textfarben haben.

```js
const title = document.querySelector('.post-title');
const excerpt = document.querySelector('.post-excerpt');

function isColorApplied(element) {
  const color = getComputedStyle(element).getPropertyValue('color');
  return color && color !== 'rgba(0, 0, 0, 0)' && color !== 'transparent'; 
}

function isMarginApplied(element) {
  const margin = getComputedStyle(element).getPropertyValue('margin');
  return margin && margin !== '0px'; 
}

assert.isTrue(isColorApplied(title));
assert.isTrue(isMarginApplied(title));

assert.isTrue(isColorApplied(excerpt));
assert.isTrue(isMarginApplied(excerpt));

const styles = new __helpers.CSSHelp(document);
// Are CSS selectors used to style the elements:
const postTitle = [".post-title", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
const postExcerpt = [".post-excerpt", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
assert.isTrue(postTitle);
assert.isTrue(postExcerpt);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="blog-post-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg" alt="Blog Post Image" class="post-img">
        <div class="post-content">
            <h2 class="post-title">Learn Web Development in 2024</h2>
            <p class="post-excerpt">Stay ahead of the curve with the latest trends in web development. Discover what's new and exciting in 2024</p>
            <a href="#" class="read-more">Read More</a>
        </div>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #f0f0f0;
    font-family: Arial, sans-serif;
}

.blog-post-card {
    background-color: white;
    border-radius: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    overflow: hidden;
    width: 350px;
    margin: 20px;
    text-align: center;
}

.post-img {
    width: 100%;
    height: auto;
    border-bottom: 5px solid #333;
}

.post-content {
    padding: 20px;
}

.post-title {
    color: #333;
    margin: 0 0 10px;
}

.post-excerpt {
    color: #667;
    margin: 10px 0;
}

.read-more {
    display: inline-block;
    margin-top: 10px;
    padding: 10px 15px;
    background-color: #333;
    color: white;
    text-decoration: none;
    border-radius: 5px;
}

.read-more:hover {
    background-color: #555;
}
```
