---
id: a2f1d72d9b908d0bd72bb9f6
title: 创建一个人员对象
challengeType: 1
forumTopicId: 16020
dashedName: make-a-person
---

# --description--

用以下方法填充对象构造函数：

```js
getFirstName()
getLastName()
getFullName()
setFirstName(first)
setLastName(last)
setFullName(first, last)
```

运行测试以查看每个方法的预期输出。 这些方法必须是与对象互动的仅有可行方法。 每一个测试会以 `new Person('Bob', 'Ross')` 声明一个新的 `Person` 实例。

# --hints--

你不应该修改函数的签名

```js
assert.match(code, /const\s+Person\s*=\s*function\s*\(\s*first\s*,\s*last\s*\)\s*{/);
```

你不应该重新设置参数 `first`

```js
assert.notMatch(code, /\bfirst\s*=\s*/);
```

你不应该重新设置参数 `last`

```js
assert.notMatch(code, /\blast\s*=\s*/);
```

不应添加任何属性。`Object.keys(Person).length` 应始终返回 6。

```js
const _person = new Person('Bob', 'Ross');
_person.setFirstName('Haskell');
_person.setLastName('Curry');
_person.setFullName('John', 'Smith');
assert.lengthOf(Object.keys(_person), 6);
```

您应该能够实例化你的 `Person` 对象。

```js
const _person = new Person('Bob', 'Ross');
assert.instanceOf(_person, Person);
```

你的 `Person` 对象不应具有 `firstName` 属性。

```js
const _person = new Person('Bob', 'Ross');
assert.notProperty(_person, 'firstName');
```

你的 `Person` 对象不应该包含 `lastName` 属性。

```js
const _person = new Person('Bob', 'Ross');
assert.notProperty(_person, 'lastName');
```

`.getFirstName()` 方法应该返回字符串 `Bob`。

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getFirstName(), 'Bob');
```

`.getLastName()` 应该返回字符串 `Ross`.

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getLastName(), 'Ross');
```

`.getFullName()` 方法应该返回字符串 `Bob Ross`.

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getFullName(), 'Bob Ross');
```

调用 `.setFirstName('Haskell')`之后，再执行`.getFullName()` 方法返回字符串 `Haskell Ross`。

```js
const _person = new Person('Bob', 'Ross');
_person.setFirstName('Haskell');
assert.strictEqual(_person.getFullName(), 'Haskell Ross');
```

调用 `.setLastName('Curry')` 之后，执行 `.getFullName()` 方法返回字符串 `Bob Curry`。

```js
const _person = new Person('Bob', 'Ross');
_person.setLastName('Curry');
assert.strictEqual(_person.getFullName(), 'Bob Curry');
```

调用 `.setFullName('Haskell', 'Curry')` 之后，执行 `.getFullName()` 方法返回字符串 `Haskell Ross`。

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getFullName(), 'Haskell Curry');
```

调用 `.setFullName('Haskell', 'Curry')` 之后，执行 `.getFullName()` 方法返回字符串 `Haskell`。

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getFirstName(), 'Haskell');
```

调用 `.setFullName('Haskell', 'Curry')` 之后，执行 `.getLastName()` 方法返回字符串 `Curry`。

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getLastName(), 'Curry');
```

调用 `.setFullName('Emily Martinez', 'de la Rosa')` 之后，执行 `.getFullName()` 方法返回字符串 `Emily Martinez de la Rosa` after calling <0>。

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getFullName(), 'Emily Martinez de la Rosa');
```

调用 `.setFullName('Emily Martinez', 'de la Rosa')` 之后，执行 `.getFirstName()` 方法返回字符串 `Emily Martinez`。

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getFirstName(), 'Emily Martinez');
```

调用 `.setFullName('Emily Martinez', 'de la Rosa')` 之后，执行 `.getLastName()` 方法返回字符串 `de la Rosa`。

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getLastName(), 'de la Rosa');
```

# --seed--

## --seed-contents--

```js
const Person = function(first, last) {
  this.getFullName = function() {
    return "";
  };
  return "";
};
```

# --solutions--

```js
const Person = function(first, last) {
  let firstName = first;
  let lastName = last;

  this.getFirstName = function(){
    return firstName;
  };

  this.getLastName = function(){
    return lastName;
  };

  this.getFullName = function(){
    return firstName + " " + lastName;
  };

  this.setFirstName = function(str){
    firstName = str;
  };


  this.setLastName = function(str){
    lastName = str;
  };

  this.setFullName = function(first, last){
    firstName = first;
    lastName = last;
  };
};
```
