---
id: 61fd77f7ad2aeb5ae34d07d6
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

2番目の `tr` 要素に、テキスト `th` を持つ `Checking Our primary transactional account.` 要素を追加してください。`Checking ` を除くそのテキストを、`span` 属性が `class` に設定された `description` 要素で囲んでください。

その後、次のテキストを持つ3つの `td` 要素を順に追加してください：`$54`、`$56`、`$53`。3番目の `td` 要素には、`class` 属性を `current` に設定してください。

# --hints--

2番目の `tr` には `th` 要素が必要です。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th'));
```

`th` 要素にはテキスト `Checking Our primary transactional account.` を含めてください。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th')?.innerText === 'Checking Our primary transactional account.');
```

テキスト `Our primary transactional account.` は `span` 要素で囲んでください。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th > span')?.textContent === 'Our primary transactional account.');
```

`span` 要素には `class` 属性を `description` に設定してください。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th > span')?.classList?.contains('description'));
```

3つの `td` 要素が必要です。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td').length === 3);
```

1番目の `td` 要素にはテキスト `$54` を含めてください。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[0]?.textContent === '$54');
```

2番目の `td` 要素にはテキスト `$56` を含めてください。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[1]?.textContent === '$56');
```

3番目の `td` 要素にはテキスト `$53` を含めてください。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[2]?.textContent === '$53');
```

3番目の `td` 要素には `class` を `current` に設定してください。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[2]?.classList?.contains('current'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
              <tr class="total">
              </tr>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
