---
id: 61fd75ea7f663457612dba02
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

新しい各`th`要素の中に、`span`が`class`に設定された`sr-only year`要素をネストしてください。次のテキストを順番に入れてください：`2019`、`2020`、`2021`。

3番目の`th`要素には、`class`属性を`current`に設定してください。

`td`要素は空のままにしてください。この要素は、テーブルが4列のレイアウトになることと、ヘッダーが正しい列に関連付けられることを保証するためだけに存在します。

# --hints--

新しい各`th`要素には`span`要素が必要です。

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.length === 1);
  assert(th?.children?.[0]?.localName === 'span');
});
```

新しい各`span`要素には、`class`属性を`sr-only year`に設定してください。

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.[0]?.classList?.contains('sr-only'));
  assert(th?.children?.[0]?.classList?.contains('year'));
});
```

最初の`span`要素にはテキスト`2019`を入れてください。

```js
assert(document.querySelectorAll('th')?.[0]?.children?.[0]?.textContent === '2019');
```

2番目の`span`要素にはテキスト`2020`を入れてください。

```js
assert(document.querySelectorAll('th')?.[1]?.children?.[0]?.textContent === '2020');
```

3番目の`span`要素にはテキスト`2021`を入れてください。

```js
assert(document.querySelectorAll('th')?.[2]?.children?.[0]?.textContent === '2021');
```

3番目の`th`要素には、`class`を`current`に設定してください。

```js
assert(document.querySelector('table')?.querySelectorAll('th')?.[2]?.classList?.contains('current'));
```

`td`要素は空のままにしてください。

```js
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.textContent === '');
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.children?.length === 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
              <tr>
                <td></td>
                <th></th>
                <th></th>
                <th></th>
              </tr>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
