---
id: 61fd6ab779390f49148773bb
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

`h1` 要素の下に `div` 要素を作成してください。`id` 属性を `years` に設定します。この特定の要素はスクリーンリーダーから隠したいので、`aria-hidden` 属性を `true` に設定してください。

# --hints--

`div` 要素の後に新しい `h1` 要素を作成する必要があります。

```js
assert(document.querySelector('h1')?.nextElementSibling?.localName === 'div');
```

新しい `div` 要素には `id` 属性を `years` に設定してください。

```js
assert(document.querySelector('div')?.getAttribute('id') === 'years');
```

新しい `div` 要素には `aria-hidden` 属性を `true` に設定してください。

```js
assert(document.querySelector('div')?.getAttribute('aria-hidden') === 'true');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
