---
id: 66bba6fff611169359d9d36a
title: Crie uma página de e-mail confidencial
challengeType: 25
dashedName: build-a-confidential-email-page
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**Histórias de usuário:**

1. Você deve ter um elemento `main` com um `id` de `email`.
1. O elemento `#email` deve ter um `padding` de `50px`, uma margem superior de `50px`, uma `width` de `500px` e um `border` de `2px`.
1. A largura total do elemento `#email`, incluindo paddings e bordas, deve ser `500px`.
1. Você deve ter dois elementos `div`, um com um id de `confidential` e o outro com um id de `top-secret`, dentro do elemento `#email`.
1. Os elementos `#confidential` e `#top-secret` devem ter um `display` de `inline-block`.
1. Os elementos `#confidential` e `#top-secret` devem ter um `padding`, uma margem esquerda e um `border`.
1. The `#confidential` element should have the text `CONFIDENTIAL`.
1. The `#top-secret` element should have the text `TOP SECRET`.
1. Os elementos `#confidential` e `#top-secret` devem ser rotacionados usando um transform do CSS.
1. Você deve ter pelo menos três elementos de parágrafo dentro do elemento `#email`.
1. Você deve ter pelo menos três elementos `span` com uma classe de `blurred` dentro de seus elementos de parágrafo.
1. Você deve ter um seletor `blurred` que desfoca o elemento `3px` usando um filtro do CSS.

# --hints--

Você deve ter um elemento `main` com um `id` de `email`.

```js
assert.exists(document.querySelector('main#email'));
```

Você deve ter um seletor `#email` que define o `padding` dos elementos em todos os lados como `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('padding'), '50px');
```

Você deve ter um seletor `#email` que define a `margin-top` dos elementos como `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('margin-top'), '50px');
```

Você deve ter um seletor `#email` que define a `width` dos elementos como `500px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('width'), '500px');
```

O elemento `#email` deve ter uma borda de `2px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('border-width'), '2px');
```

O elemento `#email` deve ter um `box-sizing` de `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.boxSizing, 'border-box');
```

Você deve ter um elemento `div` com um `id` de `confidential` dentro do elemento `#email`.

```js
assert.exists(document.querySelector('#email > div#confidential'));
```

Você deve ter um elemento `div` com um `id` de `top-secret` dentro do elemento `#email`.

```js
assert.exists(document.querySelector('#email > div#top-secret'));
```

O elemento `#confidential` deve ter o `display` definido como `inline-block`.

```js
assert.equal(window.getComputedStyle(document.querySelector('#confidential'))?.display, 'inline-block');
```

O elemento `#top-secret` deve ter o `display` definido como `inline-block`.

```js
assert.equal(window.getComputedStyle(document.querySelector('#top-secret'))?.display, 'inline-block');
```

O elemento `#confidential` deve ter um `padding` em todos os lados.

```js
const style = window.getComputedStyle(document.querySelector('#confidential'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

O elemento `#top-secret` deve ter um `padding` em todos os lados.

```js
const style = window.getComputedStyle(document.querySelector('#top-secret'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

O elemento `#confidential` deve ter uma margem esquerda.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.marginLeft, '0px');
```

O elemento `#top-secret` deve ter uma margem esquerda.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.marginLeft, '0px');
```

O elemento `#confidential` deve ter uma `border`.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.borderWidth, '0px');
```

O elemento `#top-secret` deve ter uma `border`.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.borderWidth, '0px');
```

Your `#confidential` element's text should be `CONFIDENTIAL`. You have either omitted the text, or have a typo.

```js
assert.equal(document.querySelector('#confidential').innerText.toLowerCase(), 'confidential' );
```

Your `#top-secret` element's text should be `TOP SECRET`. You have either omitted the text, or have a typo.

```js
assert.equal(document.querySelector('#top-secret').innerText.toLowerCase(), 'top secret' );
```

O elemento `#confidential` deve usar um `transform` em CSS para rotacionar o elemento.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.transform, 'none');
```

O elemento `#top-secret` deve usar um `transform` do CSS para rotacionar o elemento.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.transform, 'none');
```

Você deve ter pelo menos três elementos de parágrafo dentro do elemento `#email`.

```js
assert.isAtLeast(document.querySelectorAll('#email > p').length, 3);
```

Você deve ter pelo menos três elementos `span` com uma classe de `blurred` dentro dos parágrafos.

```js
assert.isAtLeast(document.querySelectorAll('p > span.blurred').length, 3);
```

Os elementos `.blurred` não devem estar vazios.

```js
const els = document.querySelectorAll('.blurred');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.isAtLeast(el.innerText.length , 1))
```

Você deve ter um seletor `.blurred` que define o `filter` dos elementos como `blur(3px)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.blurred')?.filter, 'blur(3px)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Confidential Email</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Secret Marshmallow Mission</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="email">
    <div id="confidential">CONFIDENTIAL</div>
    <p>Dear Agent <span class="blurred">S'more,</span></p>
    <p>We have an emergency. The secret formula for our <span class="blurred">Mega Marshmallow</span> has been
      compromised. This formula is what makes our marshmallows the fluffiest and most delicious.</p>
    <p>We suspect that <span class="blurred">Professor Puff</span> is behind this. He has taken the formula to his
      hidden laboratory. Your mission is to infiltrate the lab and secure the formula before it's too late.</p>
    <p>Be sure to keep the lab's location confidential. Any leak of this information could jeopardize the entire
      operation.</p>
    <div id="top-secret">TOP SECRET</div>
  </main>
</body>

</html>
```

```css
#email {
  background-color: SeaShell;
  margin: 50px auto;
  padding: 50px;
  width: 500px;
  border: 2px solid black;
  box-shadow: 5px 3px 3px gray;
  box-sizing: border-box;
}

#confidential {
  display: inline-block;
  margin-left: 100px;
  padding: 10px;
  text-align: center;
  font-size: 30px;
  transform: rotate(-20deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

#top-secret {
  margin-left: 200px;
  display: inline-block;
  padding: 10px;
  text-align: center;
  font-size: 20px;
  transform: rotate(15deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

.blurred {
  filter: blur(3px);
}
```
