---
id: 612ea4c4993aba52ab4aa32e
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

이제 이전에 설정한 가상 선택자를 사용할 차례입니다. 검은 건반을 만들려면 새로운 `.key.black--key::after` 선택자를 추가하세요. 이 선택자는 `key`와 `black--key` 클래스를 모두 가진 요소를 대상으로 하며, HTML에서 이 요소들 뒤에 생성된 가상 요소를 선택합니다.

새 선택자에서 `background-color`를 `#1d1e22`로 설정하세요. 또한 `content` 속성을 `""`로 설정하세요. 이렇게 하면 가상 요소가 비어 있게 됩니다.

`content` 속성은 요소의 내용을 설정하거나 재정의하는 데 사용됩니다. 기본적으로 `::before`와 `::after` 선택자에 의해 생성된 가상 요소는 비어 있어서 페이지에 렌더링되지 않습니다. `content` 속성을 빈 문자열 `""`로 설정하면 가상 요소가 렌더링되면서도 비어 있는 상태로 나타납니다.

실험해보고 싶다면 `background-color` 속성을 제거하고 `content` 속성에 `"♥"` 같은 다른 값을 설정해 보세요. 테스트가 통과하도록 작업이 끝나면 변경 사항을 반드시 되돌리세요.

# --hints--

`.key.black--key::after` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

`.key.black--key::after` 선택자에 `background-color` 속성이 `#1d1e22`로 설정되어 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor, 'rgb(29, 30, 34)');
```

`.key.black--key::after` 선택자에 `content` 속성이 `""`로 설정되어 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content, '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
