---
id: 612ec19d5268da7074941f84
title: Step 30
challengeType: 0
dashedName: step-30
---

# --description--

논리 연산자는 더 복잡한 미디어 쿼리를 구성하는 데 사용할 수 있습니다. `and` 논리 연산자는 두 가지 미디어 조건을 쿼리하는 데 사용됩니다.

예를 들어, 화면 너비가 500px 이상 1000px 이하인 경우를 대상으로 하는 미디어 쿼리는 다음과 같습니다.

```css
@media (min-width: 500px) and (max-width: 1000px){

}
```

브라우저 창 너비가 `@media`보다 크고 `769px`보다 작을 때 적용할 또 다른 `1199px` 규칙을 추가하세요.

# --hints--

새로운 `@media` 쿼리를 추가해야 합니다.

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('media'), 2);
```

`@media` 쿼리는 `min-width`인 `769px`와 `max-width`인 `1199px`를 가져야 합니다.

```js
const mediaText = new __helpers.CSSHelp(document).getCSSRules('media')[1]?.media?.mediaText;
assert.oneOf(mediaText, ['(max-width: 1199px) and (min-width: 769px)', '(min-width: 769px) and (max-width: 1199px)']);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <img class="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg" alt="freeCodeCamp Logo" />
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
  position: relative;
  border-radius: 10px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
  overflow: hidden;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
  border-radius: 0 0 3px 3px;
}

.key.black--key::after {
  background-color: #1d1e22;
  content: "";
  position: absolute;
  left: -18px;
  width: 32px;
  height: 100px;
  border-radius: 0 0 3px 3px;
}

.logo {
  width: 200px;
  position: absolute;
  top: 23px;
}

@media (max-width: 768px) {
  #piano {
    width: 358px;
  }

  .keys {
    width: 318px;
  }

  .logo {
    width: 150px;
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
