---
id: 612eb4893b63c75bb9251ddf
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

`.logo` 선택자를 만들어 로고 스타일링을 시작하세요. `width`을 `200px`로, `position`을 `absolute`로, `top`를 `23px`로 설정하세요.

# --hints--

`.logo` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.logo'));
```

`.logo` 선택자에 `width` 속성을 `200px`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.logo')?.width, '200px');
```

`.logo` 선택자에 `position` 속성을 `absolute`로 설정해야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.logo')?.position, 'absolute');
```

`.logo` 선택자에 `top` 속성을 `23px`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.logo')?.top, '23px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <img class="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg" alt="freeCodeCamp Logo" />
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

.key.black--key::after {
  background-color: #1d1e22;
  content: "";
  position: absolute;
  left: -18px;
  width: 32px;
  height: 100px;
}

--fcc-editable-region--

--fcc-editable-region--
```
