---
id: 612e89562043183c86df287c
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

브라우저는 특정 요소에 기본 여백과 패딩 값을 자주 적용합니다. 피아노 레이아웃이 올바르게 표시되도록 하려면 박스 모델을 초기화해야 합니다.

CSS 파일에 `html` 규칙 선택자를 추가하고 `box-sizing` 속성을 `border-box`로 설정하세요.

# --hints--

`html` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

`html` 선택자에 `box-sizing` 속성이 `border-box`로 설정되어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
