---
id: 612e78af05201622d4bab8aa
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

`div` 요소 안에 `body`로 설정된 `id`를 가진 `piano` 요소를 만드세요.

# --hints--

새로운 `div` 요소를 만들어야 합니다.

```js
const div = document.querySelector('div');
assert.exists(div);
```

`div`는 `body` 요소 안에 있어야 합니다.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

`div`는 `id`로 설정된 `piano`를 가져야 합니다.

```js
const div = document.querySelector('div');
assert.equal(div?.getAttribute('id'), 'piano');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css

```
