---
id: 63ec3287b182ec0efe8a3135
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

같은 자료 구조를 따라, 이 표에 있는 제품들을 순서대로 `products` 배열에 추가하세요. 각 제품마다 `id`을 1씩 증가시키면서 세세요.

| `name`                         | `price` | `category`  |
| ------------------------------ | ------- | ----------- |
| `French Macaron`               | `3.99`  | `Macaron`   |
| `Pumpkin Cupcake`              | `3.99`  | `Cupcake`   |
| `Chocolate Cupcake`            | `5.99`  | `Cupcake`   |
| `Chocolate Pretzels (4 Pack)`  | `10.99` | `Pretzel`   |
| `Strawberry Ice Cream`         | `2.99`  | `Ice Cream` |
| `Chocolate Macarons (4 Pack)`  | `9.99`  | `Macaron`   |
| `Strawberry Pretzel`           | `4.99`  | `Pretzel`   |
| `Butter Pecan Ice Cream`       | `2.99`  | `Ice Cream` |
| `Rocky Road Ice Cream`         | `2.99`  | `Ice Cream` |
| `Vanilla Macarons (5 Pack)`    | `11.99` | `Macaron`   |
| `Lemon Cupcakes (4 Pack)`      | `12.99` | `Cupcake`   |

# --hints--

`products` 배열의 두 번째 객체는 `id` 속성이 숫자 `2`으로 설정되어야 합니다.

```js
assert.equal(products[1].id, 2);
```

`products` 배열의 두 번째 객체는 `name` 속성이 `French Macaron`으로 설정되어야 합니다.

```js
assert.equal(products[1].name, 'French Macaron');
```

`products` 배열의 두 번째 객체는 `price` 속성이 숫자 `3.99`으로 설정되어야 합니다.

```js
assert.equal(products[1].price, 3.99);
```

`products` 배열의 두 번째 객체는 `category` 속성이 `Macaron`으로 설정되어야 합니다.

```js
assert.equal(products[1].category, 'Macaron');
```

`products` 배열의 세 번째 객체는 `id` 속성이 숫자 `3`으로 설정되어야 합니다.

```js
assert.equal(products[2].id, 3);
```

`products` 배열의 세 번째 객체는 `name` 속성이 `Pumpkin Cupcake`으로 설정되어야 합니다.

```js
assert.equal(products[2].name, 'Pumpkin Cupcake');
```

`products` 배열의 세 번째 객체는 `price` 속성이 숫자 `3.99`으로 설정되어야 합니다.

```js
assert.equal(products[2].price, 3.99);
```

`products` 배열의 세 번째 객체는 `category` 속성이 `Cupcake`으로 설정되어야 합니다.

```js
assert.equal(products[2].category, 'Cupcake');
```

`products` 배열의 네 번째 객체는 `id` 속성이 숫자 `4`으로 설정되어야 합니다.

```js
assert.equal(products[3].id, 4);
```

`products` 배열의 네 번째 객체는 `name` 속성이 `Chocolate Cupcake`으로 설정되어야 합니다.

```js
assert.equal(products[3].name, 'Chocolate Cupcake');
```

`products` 배열의 네 번째 객체는 `price` 속성이 숫자 `5.99`으로 설정되어야 합니다.

```js
assert.equal(products[3].price, 5.99);
```

`products` 배열의 네 번째 객체는 `category` 속성이 `Cupcake`으로 설정되어야 합니다.

```js
assert.equal(products[3].category, 'Cupcake');
```

`products` 배열의 다섯 번째 객체는 `id` 속성이 숫자 `5`으로 설정되어야 합니다.

```js
assert.equal(products[4].id, 5);
```

`products` 배열의 다섯 번째 객체는 `name` 속성이 `Chocolate Pretzels (4 Pack)`으로 설정되어야 합니다.

```js
assert.equal(products[4].name, 'Chocolate Pretzels (4 Pack)');
```

`products` 배열의 다섯 번째 객체는 `price` 속성이 숫자 `10.99`으로 설정되어야 합니다.

```js
assert.equal(products[4].price, 10.99);
```

`products` 배열의 다섯 번째 객체는 `category` 속성이 `Pretzel`으로 설정되어야 합니다.

```js
assert.equal(products[4].category, 'Pretzel');
```

`products` 배열의 여섯 번째 객체는 `id` 속성이 숫자 `6`으로 설정되어야 합니다.

```js
assert.equal(products[5].id, 6);
```

`products` 배열의 여섯 번째 객체는 `name` 속성이 `Strawberry Ice Cream`으로 설정되어야 합니다.

```js
assert.equal(products[5].name, 'Strawberry Ice Cream');
```

`products` 배열의 여섯 번째 객체는 `price` 속성이 숫자 `2.99`으로 설정되어야 합니다.

```js
assert.equal(products[5].price, 2.99);
```

`products` 배열의 여섯 번째 객체는 `category` 속성이 `Ice Cream`으로 설정되어야 합니다.

```js
assert.equal(products[5].category, 'Ice Cream');
```

`products` 배열의 일곱 번째 객체는 `id` 속성이 숫자 `7`으로 설정되어야 합니다.

```js
assert.equal(products[6].id, 7);
```

`products` 배열의 일곱 번째 객체는 `name` 속성이 `Chocolate Macarons (4 Pack)`으로 설정되어야 합니다.

```js
assert.equal(products[6].name, 'Chocolate Macarons (4 Pack)');
```

`products` 배열의 일곱 번째 객체는 `price` 속성이 숫자 `9.99`으로 설정되어야 합니다.

```js
assert.equal(products[6].price, 9.99);
```

`products` 배열의 일곱 번째 객체는 `category` 속성이 `Macaron`으로 설정되어야 합니다.

```js
assert.equal(products[6].category, 'Macaron');
```

`products` 배열의 여덟 번째 객체는 `id` 속성이 숫자 `8`으로 설정되어야 합니다.

```js
assert.equal(products[7].id, 8);
```

`products` 배열의 여덟 번째 객체는 `name` 속성이 `Strawberry Pretzel`으로 설정되어야 합니다.

```js
assert.equal(products[7].name, 'Strawberry Pretzel');
```

`products` 배열의 여덟 번째 객체는 `price` 속성이 숫자 `4.99`으로 설정되어야 합니다.

```js
assert.equal(products[7].price, 4.99);
```

`products` 배열의 여덟 번째 객체는 `category` 속성이 `Pretzel`으로 설정되어야 합니다.

```js
assert.equal(products[7].category, 'Pretzel');
```

`products` 배열의 아홉 번째 객체는 `id` 속성이 숫자 `9`으로 설정되어야 합니다.

```js
assert.equal(products[8].id, 9);
```

`products` 배열의 아홉 번째 객체는 `name` 속성이 `Butter Pecan Ice Cream`으로 설정되어야 합니다.

```js
assert.equal(products[8].name, 'Butter Pecan Ice Cream');
```

`products` 배열의 아홉 번째 객체는 `price` 속성이 숫자 `2.99`으로 설정되어야 합니다.

```js
assert.equal(products[8].price, 2.99);
```

`products` 배열의 아홉 번째 객체는 `category` 속성이 `Ice Cream`으로 설정되어야 합니다.

```js
assert.equal(products[8].category, 'Ice Cream');
```

`products` 배열의 열 번째 객체는 `id` 속성이 숫자 `10`으로 설정되어야 합니다.

```js
assert.equal(products[9].id, 10);
```

`products` 배열의 열 번째 객체는 `name` 속성이 `Rocky Road Ice Cream`으로 설정되어야 합니다.

```js
assert.equal(products[9].name, 'Rocky Road Ice Cream');
```

`products` 배열의 열 번째 객체는 `price` 속성이 숫자 `2.99`으로 설정되어야 합니다.

```js
assert.equal(products[9].price, 2.99);
```

`products` 배열의 열 번째 객체는 `category` 속성이 `Ice Cream`으로 설정되어야 합니다.

```js
assert.equal(products[9].category, 'Ice Cream');
```

`products` 배열의 열한 번째 객체는 `id` 속성이 숫자 `11`으로 설정되어야 합니다.

```js
assert.equal(products[10].id, 11);
```

`products` 배열의 열한 번째 객체는 `name` 속성이 `Vanilla Macarons (5 Pack)`으로 설정되어야 합니다.

```js
assert.equal(products[10].name, 'Vanilla Macarons (5 Pack)');
```

`products` 배열의 열한 번째 객체는 `price` 속성이 숫자 `11.99`으로 설정되어야 합니다.

```js
assert.equal(products[10].price, 11.99);
```

`products` 배열의 열한 번째 객체는 `category` 속성이 `Macaron`으로 설정되어야 합니다.

```js
assert.equal(products[10].category, 'Macaron');
```

`products` 배열의 열두 번째 객체는 `id` 속성이 숫자 `12`으로 설정되어야 합니다.

```js
assert.equal(products[11].id, 12);
```

`products` 배열의 열두 번째 객체는 `name` 속성이 `Lemon Cupcakes (4 Pack)`으로 설정되어야 합니다.

```js
assert.equal(products[11].name, 'Lemon Cupcakes (4 Pack)');
```

`products` 배열의 열두 번째 객체는 `price` 속성이 숫자 `12.99`으로 설정되어야 합니다.

```js
assert.equal(products[11].price, 12.99);
```

`products` 배열의 열두 번째 객체는 `category` 속성이 `Cupcake`으로 설정되어야 합니다.

```js
assert.equal(products[11].category, 'Cupcake');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp Shopping Cart</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1 class="title">Desserts Page</h1>
    </header>
    <main>
      <button id="cart-btn" type="button" class="btn">
        <span id="show-hide-cart">Show</span> Cart
      </button>
      <div id="cart-container">
        <button class="btn" id="clear-cart-btn">Clear Cart</button>
        <div id="products-container"></div>
        <p>Total number of items: <span id="total-items">0</span></p>
        <p>Subtotal: <span id="subtotal">$0</span></p>
        <p>Taxes: <span id="taxes">$0</span></p>
        <p>Total: <span id="total">$0</span></p>
      </div>
      <div id="dessert-card-container"></div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --black: #000;
  --white: #fff;
  --grey: #3b3b4f;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.title {
  color: var(--light-grey);
  text-align: center;
  margin: 25px 0;
}

#dessert-card-container {
  display: flex;
  flex-direction: column;
  flex-wrap: wrap;
  align-items: center;
}

.dessert-card {
  background-color: var(--light-grey);
  padding: 15px;
  text-align: center;
  border-radius: 15px;
  margin: 20px 10px;
}

.dessert-price {
  font-size: 1.2rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 100px;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#cart-btn {
  position: absolute;
  top: 0;
  right: 0;
}

.add-to-cart-btn {
  margin: 30px auto 10px;
}

#cart-container {
  display: none;
  position: absolute;
  top: 60px;
  right: 0;
  background-color: var(--light-grey);
  width: 200px;
  height: 400px;
  border: 8px double var(--black);
  border-radius: 15px;
  text-align: center;
  font-size: 1.2rem;
  overflow-y: scroll;
}

.product {
  margin: 25px 0;
}

.product-count {
  display: inline-block;
  margin-right: 10px;
}

.product-category {
  margin: 10px 0;
}

@media (min-width: 768px) {
  #dessert-card-container {
    flex-direction: row;
  }

  .dessert-card {
    flex: 1 0 21%;
  }

  #cart-container {
    width: 300px;
  }
}
```

```js
const cartContainer = document.getElementById("cart-container");
const productsContainer = document.getElementById("products-container");
const dessertCards = document.getElementById("dessert-card-container");
const cartBtn = document.getElementById("cart-btn");
const clearCartBtn = document.getElementById("clear-cart-btn");
const totalNumberOfItems = document.getElementById("total-items");
const cartSubTotal = document.getElementById("subtotal");
const cartTaxes = document.getElementById("taxes");
const cartTotal = document.getElementById("total");
const showHideCartSpan = document.getElementById("show-hide-cart");
let isCartShowing = false;

--fcc-editable-region--
const products = [
  {
    id: 1,
    name: "Vanilla Cupcakes (6 Pack)",
    price: 12.99,
    category: "Cupcake",
  },

];
--fcc-editable-region--
```
