---
id: a6e40f1041b06c996f7b2406
title: Construye un Buscador del Primer Elemento
challengeType: 26
dashedName: build-a-first-element-finder
---

# --description--

En este laboratorio crearás una función que mire a través de una matriz y devuelva el primer elemento que pase una función de prueba (también conocida como "prueba de la verdad"). 

La función iteraría a través de la matriz y evaluaría cada elemento usando la función de prueba proporcionada. Al final, devolvería el primer elemento donde la función de prueba devuelva `true`.

Ejemplo:

```js
findElement([1, 3, 5, 8], num => num % 2 === 0) // returns 8
findElement([1, 3, 5], num => num % 2 === 0)    // returns undefined
```

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes tener una función llamada `findElement` que acepte un array y una función como argumentos.
2. La función debe devolver el primer elemento de la matriz que pase una prueba de la verdad. Esto significa que, llamando a la función `func` pasada, dado un elemento `x`, la prueba de la verdad se pasa si `func(x)` es `true`.
3. Si ningún elemento pasa la prueba. la función debería devolver `undefined`.

# --hints--

Debes tener una función `findElement`.

```js
assert.isFunction(findElement);
```

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })` debe devolver `8`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })` debe devolver `undefined`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

`findElement([1, 2, 3, 4], function(num) { return num > 2; })` debería devolver `3`.

```js
assert.strictEqual(
  findElement([1, 2, 3, 4], function (num) {
    return num > 2;
  }),
  3
);
```

`findElement(["hello", "world", "javascript"], function(str) { return str.length > 5; })` debería devolver `"javascript"`.

```js
assert.strictEqual(
  findElement(["hello", "world", "javascript"], function (str) {
    return str.length > 5;
  }),
  "javascript"
);
```

`findElement(["cat", "dog", "bird"], function(str) { return str.length > 10; })` debería devolver `undefined`.

```js
assert.strictEqual(
  findElement(["cat", "dog", "bird"], function (str) {
    return str.length > 10;
  }),
  undefined
);
```

`findElement([2, 4, 6, 8], function(num) { return num % 2 === 0; })` debe devolver `2`.

```js
assert.strictEqual(
  findElement([2, 4, 6, 8], function (num) {
    return num % 2 === 0;
  }),
  2
);
```

`findElement([], function(num) { return num > 0; })` debería devolver `undefined`.

```js
assert.strictEqual(
  findElement([], function (num) {
    return num > 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
