---
id: 612ea4c4993aba52ab4aa32e
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Jetzt ist es Zeit, die zuvor eingerichteten Pseudo-Selektoren zu verwenden. Um die schwarzen Tasten zu erstellen, fügen Sie einen neuen `.key.black--key::after`-Selektor hinzu. Dieser zielt auf Elemente mit den Klassen `key` und `black--key` ab und wählt das Pseudoelement aus, das nach diesen Elementen im HTML erzeugt wird.

Legen Sie im neuen Selektor die `background-color`-Eigenschaft auf `#1d1e22` fest. Setzen Sie außerdem die `content`-Eigenschaft auf `""`. Dadurch werden die Pseudoelemente leer.

Die `content`-Eigenschaft wird verwendet, um den Inhalt eines Elements festzulegen oder zu überschreiben. Standardmäßig sind die Pseudoelemente, die durch die `::before`- und `::after`-Selektoren erzeugt werden, leer, was bedeutet, dass sie auf der Seite nicht gerendert werden. Indem Sie die `content`-Eigenschaft auf einen leeren String `""` setzen, stellen Sie sicher, dass die Pseudoelemente gerendert werden, während sie weiterhin leer erscheinen.

Wenn Sie experimentieren möchten, versuchen Sie, die `background-color`-Eigenschaft zu entfernen und verschiedene Werte für die `content`-Eigenschaft zu setzen, wie zum Beispiel `"♥"`. Denken Sie daran, diese Änderungen rückgängig zu machen, wenn Sie fertig sind, damit die Tests bestehen.

# --hints--

Sie sollten einen `.key.black--key::after`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

Ihr `.key.black--key::after`-Selektor sollte eine `background-color`-Eigenschaft mit dem Wert `#1d1e22` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor, 'rgb(29, 30, 34)');
```

Ihr `.key.black--key::after`-Selektor sollte eine `content`-Eigenschaft mit dem Wert `""` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content, '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
