---
id: 612e78af05201622d4bab8aa
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Erstellen Sie ein `div`-Element innerhalb Ihres `body`-Elements mit dem `id`, der auf `piano` gesetzt ist.

# --hints--

Sie sollten ein neues `div`-Element erstellen.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Ihr `div` sollte sich innerhalb Ihres `body`-Elements befinden.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

Ihr `div` sollte den `id` auf `piano` gesetzt haben.

```js
const div = document.querySelector('div');
assert.equal(div?.getAttribute('id'), 'piano');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css

```
