---
id: 612e83ec2eca1e370f830511
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Fügen Sie ein `link`-Element innerhalb Ihres `head`-Elements hinzu. Setzen Sie sein `rel`-Attribut auf `stylesheet` und sein `href`-Attribut auf `styles.css`.

# --hints--

Ihr Code sollte ein `link`-Element enthalten.

```js
assert.isNotNull(document.querySelector('link'));
```

Sie sollten Ihre vorhandenen `head`-Tags nicht ändern. Stellen Sie sicher, dass Sie das schließende Tag nicht gelöscht haben.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

Sie sollten ein `link`-Element haben.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Ihr `link`-Element sollte sich innerhalb Ihres `head`-Elements befinden.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Ihr `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Ihr `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    --fcc-editable-region--
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    --fcc-editable-region--
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css

```
