---
id: 612e804c54d5e7308d7ebe56
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Fügen Sie innerhalb Ihres `.keys`-Elements sieben `div`-Elemente hinzu, jeweils mit der Klasse `key`.

# --hints--

Sie sollten sieben neue `div`-Elemente erstellen.

```js
const divDivDiv = document.querySelectorAll('div');
assert.lengthOf(divDivDiv, 9);
```

Ihre sieben neuen `div`-Elemente sollten sich innerhalb Ihres `.keys`-Elements befinden.

```js
const keys = document.querySelector('.keys');
assert.lengthOf([...keys?.children], 7);
[...keys?.children].forEach(
  child => assert.equal(child.tagName, 'DIV')
)
```

Ihre sieben neuen `div`-Elemente sollten alle die `class` auf `key` gesetzt haben.

```js
const keys = document.querySelector('.keys');
[...keys?.children].forEach(child => assert.isTrue(child?.classList.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```

```css

```
