---
id: 612e8279827a28352ce83a72
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Kopieren Sie nun die Gruppe von sieben `.key`-Elementen und fügen Sie zwei weitere Gruppen in das `.keys`-Div ein.

# --hints--

Sie sollten 21 `.key`-Elemente haben.

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 21);
```

Sie sollten 15 `.black--key`-Elemente haben.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 15);
```

Alle `.key`-Elemente sollten sich innerhalb Ihres `.keys`-Elements befinden.

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.key'), 21);
```

Alle `.black--key`-Elemente sollten sich innerhalb Ihres `.keys`-Elements befinden.

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.black--key'), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

--fcc-editable-region--
        
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```

```css

```
