---
id: 612e7d1c29fb872d6384379c
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Verschachteln Sie ein zweites `div` innerhalb Ihres bestehenden `div` und setzen Sie das `class` auf `keys`.

# --hints--

Sie sollten ein zweites `div` Element erstellen.

```js
const divDiv = document.querySelectorAll('div');
assert.lengthOf(divDiv, 2);
```

Ihr neues `div` Element sollte sich innerhalb Ihres bestehenden `div` Elements befinden.

```js
const div = document.querySelector('div');
assert.lengthOf(div?.children, 1);
assert.equal(div?.children?.[0]?.localName, 'div');
```

Ihr neues `div` Element sollte das `class` auf `keys` gesetzt haben.

```js
const div = document.querySelector('div');
assert.equal(div?.children?.[0]?.className, 'keys');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
  --fcc-editable-region--
      
  --fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
