---
id: 587d78ac367417b2b2512af6
title: Alinea elementos mediante la propiedad justify-content
challengeType: 0
forumTopicId: 301102
dashedName: align-elements-using-the-justify-content-property
---

# --description--

Algunas veces los elementos flexibles dentro de un contenedor flexible no llenan todo el espacio del contenedor. Es común querer indicarle al CSS cómo alinear y espaciar los elementos flexibles de una determinada manera. Afortunadamente, la propiedad `justify-content` tiene varias opciones para hacer esto. Pero primero, hay que entender alguna terminología importante antes de revisar dichas opciones.

<a href="https://www.freecodecamp.org/news/flexbox-the-ultimate-css-flex-cheatsheet/" target="_blank" rel="noopener noreferrer nofollow">Para más información sobre las propiedades de flex-box, lee aquí</a>

Recuerda que establecer un contenedor flexible como fila coloca los elementos flexibles uno al lado del otro de izquierda a derecha. Un contenedor flexible establecido como columna coloca los elementos flexibles apilados verticalmente de arriba a abajo. Para cada uno, la dirección en la que están dispuestos los elementos flexibles se llama el **eje principal**. Para una fila, esta es una línea horizontal que recorta cada elemento. Y para una columna, el eje principal es una línea vertical a través de los elementos.

Hay varias opciones para espaciar los elementos flexibles a lo largo de la línea que representa el eje principal. Uno de los más utilizados es `justify-content: center;`, el cual alinea hacia el centro todos los elementos flexibles dentro del contenedor flexible. Otras opciones incluyen:

<ul><li><code>flex-start</code>: alinea los elementos al inicio del contenedor flexible. Para una fila, esto empuja los elementos hacia la izquierda del contenedor. Para una columna, los empuja hacia la parte superior del contenedor. Esta es la alineación predeterminada si no se especifica <code>justify-content</code>.</li><li><code>flex-end</code>: alinea los elementos al final del contenedor flexible. Para una fila, esto empuja los elementos hacia la derecha del contenedor. Para una columna, los empuja hacia la parte inferior del contenedor.</li><li><code>space-between</code>: alinea los elementos en el centro del eje principal, con espacio extra colocado entre los elementos. El primer y último elemento se empujan hasta el borde del contenedor flexible. Por ejemplo, en una fila el primer elemento está contra el lado izquierdo del contenedor, el último elemento está contra el lado derecho del contenedor, y el espacio restante se distribuye uniformemente entre los demás elementos.</li><li><code>space-around</code>: similar a <code>space-between</code> pero el primer y último elemento no están fijados a los bordes del contenedor, el espacio se distribuye alrededor de todos los elementos con medio espacio en cada extremo del contenedor flexible.</li><li><code>space-evenly</code>: distribuye el espacio de manera uniforme entre los elementos flexibles con un espacio completo en cada extremo del contenedor flexible.</li></ul>

# --instructions--

Un ejemplo ayuda a mostrar esta propiedad en acción. Agrega la propiedad CSS `justify-content` al elemento `#box-container` y asígnale un valor de `center`.

**Bonus**Prueba las otras opciones para la propiedad `justify-content` en el editor de código para ver sus diferencias. Pero ten en cuenta que un valor de `center` es el único que pasará este desafío.

# --hints--

El elemento `#box-container` debe tener una propiedad `justify-content` establecida en un valor de `center`.

```js
const boxContainer = document.querySelector('#box-container');
const justifyDirection =
  window.getComputedStyle(boxContainer)['justify-content'];
assert.strictEqual(justifyDirection, 'center');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 100%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 100%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;
    justify-content: center;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 100%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 100%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
