---
id: 66b62d0ad68488dd76228d6c
title: Depura un Cambiador Aleatorio de Color de Fondo
challengeType: 25
dashedName: debug-a-random-background-color-changer
demoType: onClick
---

# --description--

Camperbot está aprendiendo JavaScript y ha intentado construir su propio Cambiador de Color de Fondo Aleatorio. Sin embargo, ha cometido algunos errores en el proceso.   

**Objetivo:** Cumple las historias de usuario a continuación y logra que todas las pruebas pasen para que el laboratorio funcione correctamente.

**Historias de usuario:**

1. Debes corregir los errores de sintaxis en el array `darkColorsArr`.
1. Debes corregir el `ReferenceError` que se lanza cuando se hace referencia al objeto `math` dentro de la función `getRandomIndex`.
1. Debes redondear `randomIndex` al número entero más cercano.
1. Debes actualizar el valor de la variable `body` para usar el nombre de método correcto y corregir el `TypeError` relacionado.
1. Debes corregir `bgHexCodeSpanElement` para que seleccione correctamente el elemento con el `id` de `bg-hex-code`.
1. Debes modificar la variable `color` dentro de `changeBackgroundColor` para que se asigne a un color aleatorio del array `darkColorsArr`.
1. Debes actualizar la variable `btn` para que apunte al elemento botón. Abre el `index.html` para ver el nombre correcto del `id` para ese elemento de botón.
1. Debes corregir la línea que añade un event listener para el evento `click` a `btn` pasando una referencia a la función `changeBackgroundColor` en el método.
1. Tu fondo debería cambiar aleatoriamente a uno de los colores en `darkColorsArr` cada vez que se haga clic en el `#btn`.

# --hints--

Debes resolver los errores de sintaxis en el array `darkColorsArr`. Recuerda que los elementos del array deben estar separados por comas.

```js
assert.isArray(darkColorsArr);
```

You should fix the capitalization error in the `math.random()` line. Debes corregir el error de capitalización en la línea de `math.random()`. 

```js
assert.match(getRandomIndex.toString(), /\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*/);
```

Deberías redondear `darkColorsArr.length * Math.random()` hacia abajo al número entero más cercano.

```js
assert.match(getRandomIndex.toString(), /Math\.floor\(\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*\)/);
```

Debes actualizar el valor de tu variable `body` para que `queryselector` use camel case.

```js
assert.match(code, /const\s+body\s*=\s*document\.querySelector\s*\(\s*("|')body\1\s*\)/);
```

Debe corregir `"bg-hex-code"` para que haga referencia correctamente al `id` de `bg-hex-code`. Recuerda que los selectores de `id` en CSS comienzan con un símbolo de almohadilla (`#`).

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

Debes llamar a la función `getRandomIndex` dentro de la línea `darkColorsArr[getRandomIndex]`.

```js
assert.match(changeBackgroundColor.toString(), /darkColorsArr\[\s*getRandomIndex\(\s*\)\s*\]/);
```

Debes corregir el nombre del `id` en la línea `"#click-btn"` para que coincida con el nombre correcto del `id` en el archivo `index.html`.

```js
assert.match(btn?.id, /btn/);
```

No debes llamar a la función `changeBackgroundColor`. En su lugar, debes pasar la referencia de la función al método `addEventListener`.

```js
assert.match(code, /btn\.addEventListener\s*\(\s*("|')click\1\s*,\s*changeBackgroundColor\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
    "#2C3E50",
    "#34495E",
    "#2C2C2C"
    "#616A6B"
    "#4A235A"
    "#2F4F4F",
    "#0E4B5A",
    "#36454F",
    "#2C3E50",
    "#800020",
  ];
  
  function getRandomIndex() {
    const randomIndex = darkColorsArr.length * math.random();
    return randomIndex;
  }
  console.log(getRandomIndex())
  
  const body = document.queryselector("body");
  const bgHexCodeSpanElement = document.querySelector("bg-hex-code");
  console.log(bgHexCodeSpanElement)
  
  function changeBackgroundColor() {
    const color = darkColorsArr[getRandomIndex];
  
    bgHexCodeSpanElement.innerText = color;
    body.style.backgroundColor = color;
  }
  const btn = document.querySelector("#click-btn");
  console.log(btn)
  
  btn.addEventListener("click", changeBackgroundColor());
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];

function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("#bg-hex-code");

function changeBackgroundColor() {
  const color = darkColorsArr[getRandomIndex()];

  bgHexCodeSpanElement.innerText = color;
  body.style.backgroundColor = color;
}
const btn = document.querySelector("#btn");

btn.addEventListener("click", changeBackgroundColor);
```
