---
id: 660ef5105b8ba095307a0e50
title: 步骤7
challengeType: 1
dashedName: step-7
---

# --description--

现在，再次在控制台打印你的变量 `character`。 你应该在控制台先看到字符串 `"Hello"`，然后是 `"World"`。

# --hints--

你应该再次使用 `console.log` 指令。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

你应该再次打印 `character` 到控制台。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*character\s*\)/g), 2);
```


新的 `console.log()` 指令应该在你重新赋值之后。

```js
const lines = __helpers.removeJSComments(code).split(/\n+/)
const reassign = lines.findIndex(l => l.match(/character\s+=\s+("|')World\1/));
const secondLog = lines.findLastIndex(l => l.match(/console\.log/));
assert.isBelow(reassign, secondLog);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";

--fcc-editable-region--
```
