---
id: 660ef19b95d3308e7dd31bb6
title: 步骤5
challengeType: 1
dashedName: step-5
---

# --description--

<dfn>console</dfn>（控制台指令）允许你打印并查看 JavaScript 的输出。 你可以使用 `console.log()` 来向控制台发送信息。 举个例子，下面这串代码将会在控制台打印 `"Naomi"`。

```js
let developer = "Naomi";
console.log(developer);
```

上面的代码通过在 `console.log()` 中的名称读取到了变量 `developer` 。 请注意，括号内的值是将要被打印的内容。

在控制台打印变量 `character` 的值。 然后点击 “Console”（控制台）按钮来查看 JavaScript 控制台。

# --hints--

`character` 变量初始化不应被更改。

```js
assert.strictEqual(character, 'Hello');
```

You should access the `console` in your code.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

You should access the `log` property of the `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

You should use parentheses to call the `.log()` method.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

You should print the `character` variable to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
