---
id: 660f0ee51d7460ce88cd248d
title: 步骤 29
challengeType: 1
dashedName: step-29
---

# --description--

你期望在控制台看到 `4` 吗？`.push()` 会在添加你给它的值后返回数组的新长度。

了解方法返回的值很重要。 花些时间尝试使用 `.push()` 和 `.pop()`。 准备就绪后，删除所有 `.push()` 和 `.pop()` 调用以及 `console.log` 语句。

# --hints--

你不应该有 `.push()` 调用。

```js
assert.notMatch(__helpers.removeJSComments(code), /\.push\(/);
```

你不应该有 `.pop()` 调用。

```js
assert.notMatch(__helpers.removeJSComments(code), /\.pop\(/);
```

你不应该有任何的 log 语句。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log/);
```

你不应该有 `popped` 变量。

```js
assert.notMatch(__helpers.removeJSComments(code), /popped/);
```

你不应该有 `pushed` 变量。

```js
assert.notMatch(__helpers.removeJSComments(code), /pushed/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let pushed = rows.push("freeCodeCamp");
console.log(pushed);
let popped = rows.pop();
console.log(popped);
console.log(rows);
--fcc-editable-region--
```
