---
id: 660f1a00ac619ddc1e259a66
title: 步骤35
challengeType: 1
dashedName: step-35
---

# --description--

`for` 循环的 <dfn>条件</dfn> 告诉循环应迭代多少次。 当 `condition` 变为假时，循环将停止。

在 JavaScript 中，布尔值可以是 `true` 或 `false`。 这些不是字符串——你稍后会了解更多有关差异的信息。

目前，你将使用 <dfn>小于</dfn> 操作符（`<`）。这允许你查看左侧的值是否小于右侧的值。例如，如果 `count` 是 `2`，则 `count < 3` 会计算为 `true`，如果 `count` 是 `4`，则计算为 `false`。

用条件替换 `"condition"` 字符串，以检查 `i` 是否小于 `count`。

# --hints--

你应该使用小于运算符。

```js
assert.match(__helpers.removeJSComments(code), /</);
```

你应该使用小于运算符来检查 `i` 是否小于 `count`。

```js
assert.match(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

你的 `for` 循环应使用 `i < count` 作为条件。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*count\s*;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; "condition"; "iteration") {

}
--fcc-editable-region--
```
