---
id: 660f47afe4c98536715d5fa4
title: 步骤 91
challengeType: 1
dashedName: step-91
---

# --description--

<dfn>严格不等式</dfn>运算符 `!==` 允许你检查两个值是否不相等，或者是否不具有相同的类型。 语法类似于相等运算符：`value !== 4`。

现在当 `continueLoop` 为 true 时，`while` 循环运行。 将 `while` 循环条件更改为 `done` 不等于 `count`。

# --hints--

你的 `while` 循环应检查 `done` 和 `count` 是否不相等。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*(?:done\s*!==\s*count|count\s*!==\s*done)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;
  rows.push(padRow(done, count));
  if (done === count) {
    continueLoop = false;
  } 
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
