---
id: 587d824c367417b2b2512c4d
title: Comparar as propriedades de dois elementos
challengeType: 2
forumTopicId: 301588
dashedName: compare-the-properties-of-two-elements
---

# --description--

Como lembrete, este projeto está sendo construído com base no seguinte projeto inicial clonado do <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

Em `tests/1_unit-tests.js`, no teste de número `#8`, no grupo de testes `Comparisons`, modifique cada `assert` para `assert.isAbove` ou para `assert.isAtMost`, de maneira que cada teste passe (seja avaliado como `true`). Não altere os argumentos passados às afirmações.

# --hints--

Todos os testes devem passar.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Você deve escolher o método correto para a primeira afirmação - `isAbove` ou `isAtMost`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isAtMost',
  '5 is at most (<=) 5'
);
```

Você deve escolher o método correto para a segunda afirmação - `isAbove` ou `isAtMost`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'isAbove', '1 is greater than 0');
```

Você deve escolher o método correto para a terceira afirmação - `isAbove` ou `isAtMost`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[2].method,
  'isAbove',
  'Math.PI = 3.14159265 is greater than 3'
);
```

Você deve escolher o método correto para a quarta afirmação - `isAbove` ou `isAtMost`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'isAtMost',
  '1 - Math.random() is > 0 and <= 1. It is atMost 1 !'
);
```

