---
id: 587d824e367417b2b2512c55
title: Testar se um objeto tem uma propriedade
challengeType: 2
forumTopicId: 301604
dashedName: test-if-an-object-has-a-property
---

# --description--

As a reminder, this project is being built upon the following starter project cloned from <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`property` afirma que o objeto factual possui uma determinada propriedade.

# --instructions--

Em `tests/1_unit-tests.js`, no teste de número `#16`, no grupo de testes `Objects`, modifique cada `assert` para `assert.property` ou para `assert.notProperty`, de maneira que cada teste passe (seja `true`). Não altere os argumentos passados às afirmações.

# --hints--

Todos os testes devem passar.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Você deve escolher o método correto para a primeira afirmação - `property` ou `notProperty`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notProperty',
  'A car has not wings'
);
```

Você deve escolher o método correto para a segunda afirmação - `property` ou `notProperty`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'property',
  'planes have engines'
);
```

Você deve escolher o método correto para a terceira afirmação - `property` ou `notProperty`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'property', 'Cars have wheels');
```

