---
id: 687d8e05a529a4838f1c961c
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Entfernen Sie als Nächstes das Schlüsselwort `pass` innerhalb der `Node` Klasse, und geben Sie der `Node` Klasse eine `__init__` Methode mit zwei Parametern, `self` und `element`.

Weisen Sie innerhalb der Methode `element` `self.element` zu. Dadurch wird der Wert gespeichert, der beim Erstellen eines Knotens übergeben wird.

# --hints--

Sie sollten eine `__init__` Methode innerhalb Ihrer `Node` Klasse definieren.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").has_function("__init__")`) 
})
```

Ihre `__init__` Methode sollte zwei Parameter haben, `self` und `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").find_function("__init__").has_args("self, element")`) 
})
```

Sie sollten `element` innerhalb der `self.element` Methode `__init__` zuweisen.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.element == "A"
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
--fcc-editable-region--
    class Node:
        pass
--fcc-editable-region--
```
