---
id: 687d8c6f066877768f7ee441
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Ersetzen Sie innerhalb Ihrer `LinkedList`-Klasse das Schlüsselwort `pass` durch eine innere Klasse namens `Node`. Diese soll einzelne Knoten darstellen, die Daten in der Liste speichern.

# --hints--

Sie sollten eine `Node`-Klasse innerhalb der `LinkedList`-Klasse erstellen. Denken Sie daran, das Schlüsselwort `pass` innerhalb der `Node`-Klasse hinzuzufügen.

```js
({
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").has_class("Node")`) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class LinkedList:
    pass
--fcc-editable-region--
```
