---
id: 687d7deea6a2a22fa8816a73
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In vorherigen Lektionen haben Sie einige gängige Datenstrukturen kennengelernt. In diesem Workshop implementieren Sie eine verkettete Liste.

Erstellen Sie zuerst eine Klasse mit dem Namen `LinkedList`, um die Struktur darzustellen, die Ihre Listenknoten hält. Fügen Sie das Schlüsselwort `pass` innerhalb der Klasse hinzu, damit der Test bestanden wird.

# --hints--

Sie sollten eine Klasse mit dem Namen `LinkedList` erstellen. Denken Sie daran, das Schlüsselwort `pass` hinzuzufügen.

```js
({
    test: () => runPython(`assert _Node(_code).has_class("LinkedList")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
