---
id: 687e344cc1993f7316af030c
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Fügen Sie innerhalb der `LinkedList`-Klasse eine `__init__`-Methode hinzu, die `self.length` auf `0` setzt. Damit wird ein Zähler initialisiert, der verfolgt, wie viele Knoten (Elemente) sich aktuell in der Liste befinden.

# --hints--

Sie sollten `self.length` innerhalb der `0`-Methode der `__init__`-Klasse auf `LinkedList` setzen.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.length == 0`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    
--fcc-editable-region--
```
