---
id: 687e366b3004a9831e4cbd41
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Fügen Sie eine Methode namens `is_empty` hinzu, die `self.length == 0` zurückgibt.

Sie verwenden diese, um zu verfolgen, wie viele Knoten sich in der Liste befinden. Mit `self.length == 0` gibt sie `True` zurück, was bedeutet, dass die Liste keine Knoten hat.

Wenn `self.length` größer als `0` ist, gibt sie `False` zurück, was bedeutet, dass die Liste mindestens einen Knoten hat.

# --hints--

Sie sollten eine `is_empty`-Methode erstellen, die `self.length == 0` zurückgibt.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.is_empty() == True
    ll.length = 1
    assert ll.is_empty() == False`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

--fcc-editable-region--
    
--fcc-editable-region--
```
