---
id: 68800a64b06e80d9a31c131b
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Setzen Sie weiterhin in der `__init__`-Methode der `LinkedList`-Klasse `self.head` auf `None`. Dies dient als Referenz auf den ersten Knoten in der Liste. Die Liste ist momentan leer, daher müssen Sie sie auf `None` setzen.


# --hints--

Sie sollten `self.head` auf `None` innerhalb der `__init__`-Methode der `LinkedList`-Klasse setzen.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.head is None`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    def __init__(self):
        self.length = 0
--fcc-editable-region--
```
