---
id: 688003f589b875af60273dd9
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Bleiben Sie weiterhin innerhalb der `__init__`-Methode und weisen Sie `None` `self.next` zu. Dadurch wird der `next`-Zeiger auf `None` gesetzt, was bedeutet, dass der neue Knoten noch auf keinen anderen Knoten in der Liste zeigt.

Später kann `next` aktualisiert werden, um auf einen anderen Knoten zu zeigen und so die Verkettung einer verketteten Liste zu bilden.

# --hints--

Innerhalb der `__init__`-Methode sollten Sie das `next`-Attribut auf `None` initialisieren.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.next is None
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
--fcc-editable-region--
        def __init__(self, element):
            self.element = element
--fcc-editable-region--
```
