---
id: 68857e0a49baf4751f737a2a
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Erstellen Sie nun nach der `is_empty`-Methode eine Methode namens `add` innerhalb der `LinkedList`-Klasse. Die neue Methode soll zwei Parameter haben, `self` und `element`.

Erstellen Sie in Ihrer neuen Methode eine Variable namens `node` und weisen Sie ihr eine neue Instanz von `Node` zu, indem Sie `self.Node(element)` verwenden. Dadurch wird ein neuer Knoten erstellt, der zur verketteten Liste hinzugefügt wird.

# --hints--

Sie sollten eine `add`-Methode innerhalb der `LinkedList`-Klasse erstellen.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").has_function("add")`
      )
    )
})
```

Ihre `add`-Methode sollte zwei Parameter haben, `self` und `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("add").has_args("self, element")`) 
})
```

Sie sollten `self.Node(element)` in der `node`-Methode `add` zuweisen.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_body().is_equivalent("node = self.Node(element)")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
