---
id: 68858a1b4e867acd87f6f7e8
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Wenn die Liste leer ist, wird der neue Knoten zum ersten Element (dem `head`). Erstellen Sie innerhalb der `add`-Methode eine `if`-Anweisung, die überprüft, ob die Liste mit `self.is_empty()` leer ist, und weisen Sie `node` `self.head` zu.

# --hints--

Sie sollten überprüfen, ob die Liste leer ist, indem Sie die `is_empty()`-Methode verwenden.

```js
({
  test: () => runPython(`
  _cond = _Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_conditions()[0]
  assert _cond.is_equivalent("self.is_empty()")
`)})
```

Sie sollten `node` innerhalb der `self.head`-Anweisung `if` zuweisen.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_body().is_equivalent("self.head = node")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
--fcc-editable-region--
        
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
