---
id: 6888fe6c745b6852ed240229
title: Schritt 18
challengeType: 20
dashedName: step-18
---

# --description--

Als Nächstes durchlaufen Sie die verkettete Liste, um den Knoten zu finden, der das zu entfernende Element enthält.

Erstellen Sie eine `while`-Schleife, die prüft, wann `current_node` nicht `None` ist und wann `current_node.element` nicht gleich `element` ist. Innerhalb der `while`-Schleife weisen Sie `current_node` `previous_node` zu.

# --hints--

Sie sollten eine `while`-Schleife innerhalb der `remove`-Methode haben.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0]
`)})
```

Ihre `while`-Schleife sollte prüfen, wann `current_node` nicht `None` ist und wann `current_node.element` nicht gleich `element` ist.

```js
({
  test: () => runPython(`
  _cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0].find_conditions()[0]
  assert _cond.is_equivalent("current_node is not None and current_node.element != element")
`)})
```

Sie sollten innerhalb der `current_node`-Schleife `previous_node` `while` zuweisen.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0]
assert _cond.is_equivalent("while current_node is not None and current_node.element != element:\\n previous_node = current_node")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
--fcc-editable-region--
        
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
