---
id: 6888f9827637322d4f309192
title: Schritt 17
challengeType: 20
dashedName: step-17
---

# --description--

Erstellen Sie innerhalb der `remove`-Methode eine Variable namens `current_node` und weisen Sie ihr den Wert von `self.head` zu.

Dies startet die Durchquerung der verketteten Liste ab dem ersten Knoten. `current_node` wird verwendet, um den Knoten zu verfolgen, der gerade untersucht wird, während Sie nach dem Knoten suchen, der das angegebene Element zum Entfernen enthält.

# --hints--

Sie sollten `self.head` innerhalb der `current_node`-Methode `remove` zuweisen.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("remove").has_stmt("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

--fcc-editable-region--
    def remove(self, element):
        previous_node = None
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
