---
id: 688ac806194cfc60ba1e1767
title: Schritt 20
challengeType: 20
dashedName: step-20
---

# --description--

Nach der `while`-Schleife erstellen Sie eine `if`-Anweisung, die überprüft, ob `current_node` `None` ist.

Wenn dies der Fall ist, wurde die verkettete Liste durchlaufen, ohne das zu entfernende Element zu finden, und die Funktion sollte die Ausführung beenden. Fügen Sie daher das Schlüsselwort `return` in den `if`-Block ein.

# --hints--

Sie sollten eine `if`-Anweisung erstellen, die überprüft, ob `current_node` `None` ist, und darin das Schlüsselwort `return` hinzufügen.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0]
assert _cond.is_equivalent("if current_node is None:\\n return")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
        while current_node is not None and current_node.element != element:
            previous_node = current_node
            current_node = current_node.next
--fcc-editable-region--
        
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
