---
id: 68b84df2648b4c96b25022f3
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Es ist Zeit, Ihre verkettete Liste bis jetzt zu testen! Erstellen Sie am Ende Ihres Codes eine Instanz von `LinkedList` mit dem Namen `my_list`. Geben Sie dann mit der `is_empty`-Methode, die Sie im letzten Schritt erstellt haben, `my_list.is_empty()` aus.

# --hints--

Sie sollten eine Instanz von `LinkedList` mit dem Namen `my_list` erstellen.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_variable("my_list").value_is_call("LinkedList")
    `)
  )
})
```

Sie sollten `my_list.is_empty()` ausgeben.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).has_call("print(my_list.is_empty())")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0

--fcc-editable-region--

--fcc-editable-region--
```
