---
id: 68b9e951e4892721632acc26
title: Schritt 15
challengeType: 20
dashedName: step-15
---

# --description--

Super! Nun ist es an der Zeit, das Hinzufügen von Elementen zu Ihrer verketteten Liste zu testen. Fügen Sie am Ende Ihres Codes, nachdem Sie eine Instanz von `LinkedList` namens `my_list` erstellt und `my_list.is_empty()` ausgegeben haben, `1` und `2` mit der `add`-Methode zur Liste hinzu.

Geben Sie anschließend `my_list.is_empty()` erneut aus, um zu bestätigen, dass die Liste nicht mehr leer ist. Zum Schluss geben Sie `my_list.length` aus, um die aktuelle Länge der verketteten Liste zu sehen.

# --hints--

Sie sollten `1` mit der `add`-Methode zur Liste hinzufügen.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(1)")`)
  )
})
```

Sie sollten `2` mit der `add`-Methode zur Liste hinzufügen.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(2)")`)
  )
})
```

Sie sollten `my_list.is_empty()` ausgeben.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[1].is_equivalent("print(my_list.is_empty())")
    `)
  )
})
```

Sie sollten `my_list.length` ausgeben.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[2].is_equivalent("print(my_list.length)")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

my_list = LinkedList()
print(my_list.is_empty())

--fcc-editable-region--

--fcc-editable-region--
```
