---
id: 6853e58d52ec64349a5760ca
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個工作坊中，你將實作最短路徑演算法。你將撰寫一個 Python 函式，計算圖中節點之間的最短路徑，並且傳回所採取的路徑。

例如，給定一個城市由不同距離的道路連接的圖形，該演算法將找到從一個城市到另一個城市的最短路徑。如果你想從 City A 旅行到 City D，該演算法可能會發現走 A ⇨ B ⇨ C ⇨ D（總共：15km）比直接走 A ⇨ D（20km）更短。

要開始，定義一個名為 `INF` 的變數並指定其值為 `float('inf')`，它表述正無限大。稍後，你會用它來顯示兩個節點之間的無限距離。

# --hints--

你應該在你的檔案頂部定義變數 `INF`。

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("INF")`));
  }
});
```

你應該將 `float('inf')` 指定給變數 `INF`。

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("INF").is_equivalent("INF = float('inf')")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
