---
id: 686253f1b051998ad4904e3e
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

你現在將創建接受三個參數的主函式（程式）：鄰接矩陣、起始節點，以及選用的標的節點。

創建一個名為 `shortest_path` 的函式（程式），該函式接受三個參數：`matrix`、`start_node` 和 `target_node`。將 `None` 指定為 `target_node` 的預設值。

`target_node` 參數的預設值為 `None`，表示如果未指定標的節點，函式（程式）應該計算從起始節點到圖中所有其他節點的最短路徑。

暫時在函式（程式）主體內添加一個 `pass` 述語。

# --hints--

你應該創建一個名為 `shortest_path` 的函式（程式）。

```js
({
    test: () => {
      assert(runPython(`_Node(_code).has_function("shortest_path")`));
    }
});
```

該函式（程式）應該接受三個參數：`matrix`、`start_node` 和 `target_node`，其中 `target_node` 的預設值為 `None`。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_args("matrix, start_node, target_node=None")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--

--fcc-editable-region--
```
