---
id: 686251fb49ebc489ba60ccb5
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

你需要創建一個 2D 列表來表述圖的鄰接矩陣。這個矩陣將用來表述圖中節點之間邊的權重。

創建一個名為 `adj_matrix` 的變數，並指定它為一個表述圖形且具有以下權重的二維列表：

```py
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

你應該創建一個名為 `adj_matrix` 的變數。

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("adj_matrix")`));
  }
});
```

`adj_matrix` 應該是一個包含所提供值的二維列表。

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("adj_matrix").is_equivalent("adj_matrix = [[0, 5, 3, INF, 11, INF], [5, 0, 1, INF, INF, 2], [3, 1, 0, 1, 5, INF], [INF, INF, 1, 0, 9, 3], [11, INF, 5, 9, 0, INF], [INF, 2, INF, 3, INF, 0]]")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
--fcc-editable-region--

--fcc-editable-region--
```
