---
id: 68625ad34aa3f58d7f3b9b07
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

你需要儲存圖中節點的數量。為此你將需要一個與鄰接矩陣長度相符的變數。

在 `shortest_path` 函式（程式）內，創建一個變數 `n` 並將其設為 `matrix` 的長度。

# --hints--

你應該將 `n` 初始化為 `matrix` 的長度。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("n").is_equivalent("n = len(matrix)")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    pass

--fcc-editable-region--
```
