---
id: 686b82af323098e7a5203bf4
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

當演算法執行時，你需要追蹤你已經造訪過的節點，以免重複處理它們。

為此，創建一個名為 `visited` 的列表，並為每個節點初始化為 `False`。

# --hints--

你應該在 `shortest_path` 函式（程式）內有一個名為 `visited` 的列表。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("visited")`
        )
      );
    }
});
```

你應該使用 `[False] * n` 為每個節點將 `visited` 列表初始化為 `False`。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("visited").is_equivalent("visited = [False] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]

--fcc-editable-region--
```
