---
id: 686b73007fc582db0132c361
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

你需要追蹤從起始節點到圖中每個其他節點的最短已知距離。

為此，創建一個名為 `distances` 的變數，並將其初始化為包含單一元素 `INF` 的列表。  


# --hints--

你應該創建一個名為 `distances` 的變數。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("distances")`
        )
      );
    }
});
```

`distances` 列表應該初始化為 `[INF]`。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF]")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)

--fcc-editable-region--
```
