---
id: 686b78bc13d518e1267448e4
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

除了追蹤距離之外，你還需要追蹤到達每個節點所採取的實際路徑。

你將創建一個列表，其中每個條目儲存到達該節點所經過的路徑。最初，每個節點的路徑只會包含它自己。  

列表推導式提供了一種簡潔的方式來創建列表。 例如：

```py
[x * 2 for x in range(3)]
```

創建一個名為 `paths` 的變數，並使用列表推導式初始化它，該推導式為 `range(n)` 中的每個 `node_no` 產生一個包含 `[node_no]` 的列表。

# --hints--

你應該有一個名為 `paths` 的變數。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("paths")`
        )
      );
    }
});
```

你應該使用列表推導式來初始化 `paths`。

```js
({  
    test: () => {  
      runPython(`  
import ast  
assert isinstance(_Node(_code).find_function("shortest_path").find_variable("paths").tree.value, ast.ListComp)  
    `);  
    }  
});
```

你的列表推導式應該使用 `node_no` 來迭代 `range(n)`。

```js
({  
  test: () => {  
    runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert len(v.find_comp_targets()) == 1  
assert v.find_comp_targets()[0].is_equivalent("node_no")  
assert len(v.find_comp_iters()) == 1  
assert v.find_comp_iters()[0].is_equivalent("range(n)")  
    `);  
  }  
});
```

你的列表推導式應該對 `range(n)` 中的每個 `node_no` 評估 `[node_no]`。  

```js  
({  
    test: () => {  
      runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert v.find_comp_expr().is_equivalent("[node_no]")  
    `);  
    }  
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0

--fcc-editable-region--
```
