---
id: 686b7848fbaa41e0c2a7efef
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

既然你已經初始化了你的 `distances` 列表，你需要更新起始節點的距離。

由於從起始節點到自身的距離永遠是 `0`，因此將 `distances` 列表中 `start_node` 索引的值設為 `0`。

# --hints--

你應該將 `distances[start_node]` 設為 `0`。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_stmt("distances[start_node] = 0")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n

--fcc-editable-region--
```
