---
id: 68760b53e250f982473e1808
title: 步驟 13
challengeType: 20
dashedName: step-13
---

# --description--

如果你剛剛添加的條件為真，表示目前節點是你迄今為止找到的最佳未訪問選項，且你需要更新你的變數以反映該狀況。

在 `if` 區塊內，將 `min_distance` 更新為 `distances[node_no]` 並將 `current` 設定為 `node_no`。

# --hints--

你應該將 `min_distance` 變數更新為 `distances[node_no]`。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("min_distance = distances[node_no]")`
        )
      );
    }
});
```

你應該將 `current` 設定為 `node_no`。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("current = node_no")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

--fcc-editable-region--
    for _ in range(n):
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            if not visited[node_no] and distances[node_no] < min_distance:  
                pass

--fcc-editable-region--
```
