---
id: 68760464eb9c2e79cc913fac
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

在此步驟中，你將添加一個迴圈，該迴圈將對圖中的每個節點執行一次。此迴圈將允許演算法在多次通過中更新距離和路徑。

創建一個執行 `n` 次的 `for` 迴圈。使用 `_` 作為迴圈變數，因為你不需要使用迭代值。  

在迴圈內，你需要透過建立兩個變數來準備選擇下一個要處理的節點： 

- 一個用來保存目前迭代中找到的最小距離
- 以及另一個用來儲存具有此最小距離的節點索引。

創建變數 `min_distance` 和 `current`，並分別將它們設為 `INF` 和 `-1`。 

# --hints--

你應該在 `shortest_path` 函式中有一個 `for` 迴圈。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0]`
        )
      );
    }
});
```

你的 `for` 迴圈應該要對 `range(n)` 進行迭代。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_iter().is_equivalent("range(n)")`
        )
      );
    }
});
```

你的 for 迴圈應該使用 `_` 作為迭代變數。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_vars().is_equivalent("_")`
        )
      );
    }
});
```

你應該在 `for` 迴圈內創建一個名為 `min_distance` 的變數並將其設為 `INF`。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("min_distance = INF")`
        )
      );
    }
});
```

你應該在 for 迴圈內創建一個名為 `current` 的變數並將其設為 `-1`。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("current = -1")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
--fcc-editable-region--

--fcc-editable-region--
```
