---
id: 6876089f7a500b7c782e5b83
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

你需要決定目前的節點是否比你已經找到的節點（如果有的話）更好。為此，你將在迴圈中添加條件述語。

該條件應該執行兩件事：

- 檢查該節點是否尚未被訪問。
- 將已知距離與目前的 `min_distance` 進行比較。

在內層的 `for` 迴圈中，為你的程式碼添加一個 `if` 述語，檢查 `node_no` 是否尚未被造訪 **且** `distances[node_no]` 是否小於 `min_distance`。

在條件區塊內將 `pass` 添加為佔位符。

# --hints--

你應該在你的內層 `for` 迴圈中添加一個 `if` 述語。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0]`
        )
      );
    }
});
```

你的 `if` 述語應該檢查節點是否未被訪問（`visited[node_no]` 為假值）且 `distances[node_no]` 小於 `min_distance`。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("not visited[node_no] and distances[node_no] < min_distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

--fcc-editable-region--
    for _ in range(n):  
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            pass

--fcc-editable-region--
```
