---
id: 68760c4c2648e7832c5eef16
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

在尋找最近未造訪節點的迴圈之後，你需要檢查是否實際找到有效的節點。

如果不存在這樣的節點，表示剩餘的節點無法從起始節點到達，演算法應該提前停止。

在與巢狀 `for` 迴圈相同的層次，添加一個 `if` 述語來檢查 `current == -1`，如果為真則跳出迴圈。

# --hints--

你應該有一個 `if` 述語，用來檢查在內部的 `for` 迴圈結束後，`current` 是否仍然是 `-1`。

```js
({
    test: () => {
      runPython(`
      cond = _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
      assert cond.is_equivalent("current == -1") or cond.is_equivalent("-1 == current")
      `);
    }
});
```

你應該使用 `break` 述語來跳出迴圈。

```js
({
    test: () => {
      runPython(`
      assert _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("break")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            if not visited[node_no] and distances[node_no] < min_distance:  
                min_distance = distances[node_no]  
                current = node_no

--fcc-editable-region--

--fcc-editable-region--
```
