---
id: 68760671d950767b055bf2f9
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

現在你需要檢查每個節點，以找出尚未被造訪且已知距離最小的節點。

為了做到這點，請在主迴圈內添加一個 `for` 迴圈。

迴圈應該迭代 `range(n)` 中的每個 `node_no`，其中 `n` 是圖中節點的數量。暫時將 `pass` 作為迴圈的主體。

# --hints--

你應該在 `shortest_path` 函式中有第二個 `for` 迴圈。

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
len(outer_loop.find_bodies()[0].find_for_loops()) >= 1
      `)
      );
    }
});
```

你的 `for` 迴圈應該以 `node_no` 作為迭代變數，對 `range(n)` 進行迭代。

```js
({
    test: () => {
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
inner_loop = outer_loop.find_bodies()[0].find_for_loops()[0]
assert inner_loop.find_for_vars().is_equivalent("node_no")
assert inner_loop.find_for_iter().is_equivalent("range(n)")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

--fcc-editable-region--
    for _ in range(n):  
        min_distance = INF  
        current = -1

--fcc-editable-region--
```
