---
id: 687774e17e0972345de2527a
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

如果在此遍歷中找到有效節點，你需要將它標記為已訪問，這樣未來的迭代中就不會再次考慮它。

在你前一步驟添加的 `if` 述語之後，將 `visited[current]` 設為 `True`。

# --hints--

你應該在最後一個 `if` 述語之後將 `visited[current]` 設為 `True`。

```js
({
    test: () => {
      runPython(`
      block = _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0]
      if_stmt = """if current == -1:
        break"""
      target_stmt = "visited[current] = True"
      assert block.is_ordered(if_stmt, target_stmt)
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            if not visited[node_no] and distances[node_no] < min_distance:  
                min_distance = distances[node_no]  
                current = node_no

--fcc-editable-region--
        if current == -1:  
            break

--fcc-editable-region--
```
