---
id: 68777816e95d4936f129819f
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

在嘗試更新到鄰居的距離之前，你需要確認該鄰居既可達又未被訪問。然後你將計算通過當前節點到達該鄰居的總距離會是多少。

在 `for` 迴圈內，添加一個檢查的 `if` 述語：

- `distance` 不相等於 `INF`（表示節點之間有邊）
- 鄰居 `node_no` 尚未被訪問。

在 `if` 區段內，創建一個名為 `new_distance` 的變數，並將 `distances[current]`（到目前節點的最短距離）與 `distance`（從目前節點到鄰居的距離）相加後指定給它。

# --hints--

你應該有一個 `if` 述語，用來檢查 `distance` 是否不相等於 `INF`，以及鄰近節點是否尚未被訪問。

```js
({
    test: () => {
      assert(
        runPython(`_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("distance != INF and not visited[node_no]") 
      `)
      );
    }
});
```

你應該在 `if` 區塊內創建一個名為 `new_distance` 的變數。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("new_distance")`
        )
      );
    }
});
```

你應該將 `new_distance` 指定為 `distances[current]` 和 `distance` 的總和。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("new_distance").is_equivalent("new_distance = distances[current] + distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
              min_distance = distances[node_no]
              current = node_no

        if current == -1:
            break
        visited[current] = True

--fcc-editable-region--
        for node_no in range(n):  
            distance = matrix[current][node_no]

--fcc-editable-region--
```
