---
id: 687b45e16f38ad5f4e4bd799
title: 第 23 步
challengeType: 20
dashedName: step-23
---

# --description--

既然你已經決定節點應該被顯示，你需要格式化它的路徑，以便能夠清楚地列印。為此你將使用產生器表達式。

產生器表達式類似於列表推導式，但它不是建立列表，而是一次產生一個值。它使用小括弧 `()` 而非方括弧 `[]`。例如：

```py
numbers = [1, 2, 3]
squared = (x**2 for x in numbers)  # Generator expression
```

在 `if` 述語之後的迴圈內，創建一個名為 `string_path` 的變數。為它指定一個產生器表達式，該表達式使用 `str()` 將 `paths[node_no]` 中的每個節點編號轉換為字串。 

產生器表達式應該對 `paths[node_no]` 中的每個節點編號 `n` 進行迭代。

# --hints--

你應該創建一個名為 `string_path` 的變數。

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
for_loop.find_bodies()[0].has_variable("string_path")
      `)
      );
    }
});
```

你應該將產生器表達式指定給 `string_path`。

```js
({
  test: () => {
      runPython(`
      import ast
      assert isinstance(_Node(_code).find_function("shortest_path").find_for_loops()[1].find_bodies()[0].find_variable("string_path").tree.value, ast.GeneratorExp)
      `)
 }
})
```

你的產生器表達式應該對 `paths[node_no]` 進行迭代。

```js
({
    test: () => {
        runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
path_var = for_loop.find_bodies()[0].find_variable("string_path")
assert path_var.find_comp_iters()[0].is_equivalent("paths[node_no]")
      `);
    }
});
```

你的產生器表達式應該使用 `n` 作為迭代變數來迭代 `paths[node_no]`。

```js
({
  test: () => {
      runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
string_path_var = for_loop.find_bodies()[0].find_variable("string_path")
assert string_path_var.find_comp_targets()[0].is_equivalent("n")
      `);
    }
});
```

你的產生器運算式應該為 `paths[node_no]` 中的每個 `n` 計算 `str(n)`。

```js
({
  test: () => {
      runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
string_path_var = for_loop.find_bodies()[0].find_variable("string_path")
assert string_path_var.find_comp_expr().is_equivalent("str(n)")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
                min_distance = distances[node_no]
                current = node_no

        if current == -1:
            break
        visited[current] = True

        for node_no in range(n):
            distance = matrix[current][node_no]
            if distance != INF and not visited[node_no]:
                new_distance = distances[current] + distance
                if new_distance < distances[node_no]:
                    distances[node_no] = new_distance
                    paths[node_no] = paths[current] + [node_no]

--fcc-editable-region--
    targets = [target_node] if target_node is not None else range(n)
    for node_no in targets:
        if node_no == start_node or distances[node_no] == INF:
            continue

--fcc-editable-region--
```
