---
id: 687b48ccb46c5e6126d30bc6
title: 第 25 步
challengeType: 20
dashedName: step-25
---

# --description--

最後，列印目前節點的結果。

你的輸出應該顯示從 `start_node` 到 `node_no` 的距離以及到達那裡所採取的路徑。

使用 `f-string` 格式化輸出，以顯示距離和完整路徑，格式如下：

```md
\n[starting node]-[node number] distance: [distance]\nPath: [path]
```

其中：

- `[starting node]` 是 `start_node`
- `[node number]` 是 `node_no`
- `[distance]` 是從 `start_node` 到 `node_no` 的距離
- `[path]` 是到達 `node_no` 所採取的路徑

# --hints--

你應該使用 `print(f'\n{start_node}-{node_no} distance: {distances[node_no]}\nPath: {path}')` 列印正確的輸出。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function('shortest_path').find_for_loops()[1].has_call('print(f"\\\\n{start_node}-{node_no} distance: {distances[node_no]}\\\\nPath: {path}")')`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
                min_distance = distances[node_no]
                current = node_no

        if current == -1:
            break
        visited[current] = True

        for node_no in range(n):
            distance = matrix[current][node_no]
            if distance != INF and not visited[node_no]:
                new_distance = distances[current] + distance
                if new_distance < distances[node_no]:
                    distances[node_no] = new_distance
                    paths[node_no] = paths[current] + [node_no]

--fcc-editable-region--
    targets = [target_node] if target_node is not None else range(n)
    for node_no in targets:
        if node_no == start_node or distances[node_no] == INF:
            continue
        string_path = (str(n) for n in paths[node_no])
        path = ' -> '.join(string_path)
        
--fcc-editable-region--
```
