---
id: 68e4c0ff9edcc92771d7549a
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

目前，`distances` 列表只有一個元素。然而，你需要為圖中的每個節點一個距離值。

在 Python 中，你可以將一個列表乘以一個整數來重複它。 例如，`[0] * 3` 會創建 `[0, 0, 0]`。

將 `[INF]` 乘以 `n`，更新 `distances` 列表以包含 `n` 個 `INF`。

# --hints--

你應該更新 `distances` 以包含 `n` 個 `INF`。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF]

--fcc-editable-region--
```
