---
id: ab306dbdcc907c7ddfc30830
title: إنشاء أداة تسطيح عميق
challengeType: 26
dashedName: create-a-deep-flattening-tool
---

# --description--

في هذا المختبر، ستنفذ خوارزمية لتسطيح المصفوفات.

تسطيح المصفوفة يعني تحويل مصفوفة مسكنة بأي عمق إلى مصفوفة أحادية البعد واحدة. تستخرج العملية كل العناصر بالترتيب، مع فك التغليف فقط للمصفوفات. تبقى الأنواع الأخرى دون تغيير.

على سبيل المثال:

| الأصلية | المسطحة |
| --- | --- |
| `[[1], [], [2, [3]]]` | `[1, 2, 3]` |
| `[1, {"foo": "bar"}, [2]]` | `[1, {"foo": "bar"}, 2]` |
| `["baz", [1, 2], {}]` | `["baz", 1, 2, {}]` |

**الهدف:** حقق قصص المستخدم أدناه واجعل كل الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تملك دالة باسم `steamrollArray`.
1. يجب أن تقبل دالة `steamrollArray` معلمة واحدة: مصفوفة مسكنة.
1. يجب أن تسطح الدالة المصفوفة المسكنة، مع الأخذ في الاعتبار مستويات التسكين المختلفة.
1. يجب ألا تستخدم الحل الخاص بك طريقتي `Array.prototype.flat()` أو `Array.prototype.flatMap()`.
1. يجب عدم استخدام المتغيرات العامة.

# --hints--

يجب أن تملك دالة `steamrollArray`.

```js
assert.isFunction(steamrollArray);
```

`steamrollArray([[["a"]], [["b"]]])` يجب أن تُرجع `["a", "b"]`.

```js
assert.deepEqual(steamrollArray([[['a']], [['b']]]), ['a', 'b']);
```

`steamrollArray([1, [2], [3, [[4]]]])` يجب أن تُرجع `[1, 2, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, [2], [3, [[4]]]]), [1, 2, 3, 4]);
```

`steamrollArray([1, [], [3, [[4]]]])` يجب أن تُرجع `[1, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, [], [3, [[4]]]]), [1, 3, 4]);
```

`steamrollArray([1, {}, [3, [[4]]]])` يجب أن تُرجع `[1, {}, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4]);
```

يجب ألا تستخدم الحل الخاص بك طريقتي `Array.prototype.flat()` أو `Array.prototype.flatMap()`.

```js
const arr = [1,2,3,[1,2,3,[1,2,3]]];

const spyFlat = __helpers.spyOn(Array.prototype, 'flat');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlat.calls);
} finally {
  spyFlat.restore();
}

const spyFlatMap = __helpers.spyOn(Array.prototype, 'flatMap');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlatMap.calls);
} finally {
  spyFlatMap.restore();
}
```

يجب ألا تستخدم متغيرات عامة.

```js
steamrollArray([1, {}, [3, [[4]]]])
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4])
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function steamrollArray(arr) {
  if (!Array.isArray(arr)) {
    return [arr];
  }
  const out = [];
  arr.forEach(function(e) {
    steamrollArray(e).forEach(function(v) {
      out.push(v);
    });
  });
  return out;
}
```
