---
id: 587d8259367417b2b2512c85
title: 實作 Selection Sort 演算法
challengeType: 26
dashedName: implement-the-selection-sort-algorithm
---

# --description--

在這個實驗中，你將實作選擇排序演算法。選擇排序的運作方式是從一個列表中選出最小值，並將它與列表中的第一個值交換。接著從第二個位置開始，選出剩餘列表中最小的值，並與第二個元素交換。它會持續遍歷列表並交換元素，直到到達列表末端。現在列表已經排序完成。選擇排序在所有情況下的時間複雜度都是二次方。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `selectionSort` 函式（程式）。
1. `selectionSort` 函式（程式）應該接受一個陣列，並傳回一個包含相同元素但從小到大排序的陣列。

# --hints--

`selectionSort` 應該是一個函式（程式）。

```js
assert.isFunction(selectionSort);
```

`selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` 應該傳回一個除了順序外未改變的陣列。

```js
assert.sameMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`selectionSort` 應該傳回一個已排序的陣列（由小到大）。

```js
assert.sameOrderedMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
)
```

`selectionSort` 不應該使用內建的 `.sort()` 方法。

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
