---
id: 6735a7370e0ae93a4577c689
title: Розробіть макет журналу
challengeType: 25
dashedName: design-a-magazine-layout
demoType: onClick
---

# --description--

У цій лабораторній роботі ви розробите макет журналу, використовуючи сітку в CSS, включно з рядками та стовпцями.

**Мета:** виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Ви повинні мати елемент `main` з класом `magazine-cover`.

2. Елемент `.magazine-cover` повинен мати елемент `header` з класом `title`, в якому вказано заголовок.

3. На сторінці має бути чотири елементи `section`. Використайте перші три елементи `section` для статей, а четвертий — для обкладинки.

4. Чотири елементи `section` повинні мати атрибут `class` з відповідними значеннями: `feature-article`, `small-article1`, `small-article2` та `cover-image`.

5. You should set the `display` property of `.magazine-cover` to `grid`.

6. Елемент `.magazine-cover` повинен мати властивість `grid-template-areas`, яка визначає макет сітки:

   - заголовок повністю охоплює перший рядок;
   - у другому рядку — актуальна стаття на два стовпці та обкладинка;
   - у третьому рядку — дві менші статті та обкладинка.

7. Елемент `.magazine-cover` повинен мати властивість `grid-template-columns`, яка ділить матеріал на три рівні частини.

8. Елемент `.magazine-cover` повинен мати властивість `grid-template-rows`. Для першого рядка задайте значення `auto`, а для двох інших — поділіть те місце, що залишилось.

9. Додайте жолоб між елементами сітки і встановіть ширину елемента `.magazine-cover` такою, щоб він займав більшу частину ширини вікна перегляду.

**Примітка:** не забудьте прив’язати таблицю стилів до HTML та застосувати CSS.

# --hints--

Ви повинні мати елемент `main`.

```js
assert.exists(document.querySelector('main'));
```

Елемент `main` повинен мати `class` зі значенням `magazine-cover`.

```js
assert.exists(document.querySelector('main.magazine-cover'));
```

В межах елемента `main` має бути елемент `header`.

```js
assert.exists(document.querySelector("main > header"));
```

Елемент `header` повинен мати `class` зі значенням `title` і містити текст.

```js
const header = document.querySelector("header.title");
assert.exists(header);
assert.isAbove(header.innerText.length, 0);
```

Ви повинні мати по одному елементу `section` для кожної статті й обкладинки з відповідними класами: `feature-article`, `small-article1`, `small-article2` та `cover-image`.

```js
const featureArticle = document.querySelector('section.feature-article');
const smallArticle1 = document.querySelector('section.small-article1');
const smallArticle2 = document.querySelector('section.small-article2');
const coverImage = document.querySelector('section.cover-image');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(coverImage);
```

Елемент `.magazine-cover` повинен мати властивість `display` зі значенням `grid`.

```js
assert.equal(getComputedStyle(document.querySelector('.magazine-cover')).display, 'grid');
```

Елемент `.magazine-cover` повинен мати властивість `grid-template-areas`, де `title` займає всі стовпці першого рядка.

```js
const magazineCoverStyle = getComputedStyle(document.querySelector('.magazine-cover'));
assert.match(magazineCoverStyle.gridTemplateAreas, /^("|')title title title\1/);
```

Другий рядок сітки повинен містити `feature article` (який займає два стовпці) та `cover image`.

```js
const magazineCoverStyle = getComputedStyle(document.querySelector('.magazine-cover'));
assert.match(magazineCoverStyle.gridTemplateAreas, /.("|')feature-article feature-article cover-image\1./);
```

Третій рядок сітки повинен містити `small-article1`, `small-article2` та `cover-image`.

```js
const magazineCoverStyle = getComputedStyle(document.querySelector('.magazine-cover'));
assert.match(magazineCoverStyle.gridTemplateAreas, /("|')small-article1 small-article2 cover-image\1$/);
```

Селектор `.magazine-cover` має використовувати властивість `grid-template-columns`, щоб поділити матеріал на три рівні частини. Ви можете використати функцію `repeat` або вручну визначити дробові частини.

```js
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.magazine-cover')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
```

Клас `.magazine-cover` повинен мати властивість `grid-template-rows` зі значенням `auto 1fr 1fr`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.magazine-cover')?.gridTemplateRows, 'auto 1fr 1fr');
```

Додайте жолоб зі значенням `10px` між елементами сітки.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.magazine-cover')?.gap, '10px');
```

Клас `.title` повинен мати область сітки з такою ж назвою.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.title')?.gridArea, 'title');
```

Клас `.feature-article` повинен мати область сітки з такою ж назвою.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.feature-article')?.gridArea, 'feature-article');
```

Клас `.cover-image` повинен мати область сітки з такою ж назвою.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cover-image')?.gridArea, 'cover-image');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Magazine Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Adventure Magazine</title>
  <link rel="stylesheet" href="styles.css">
</head>
<body>
  <main class="magazine-cover">
      <header class="title">
          <h1>Adventure Awaits</h1>
      </header>
      <section class="feature-article">
          <h2>Top 10 Exotic Destinations for 2024</h2>
          <p>Discover the most breathtaking places to visit this year, from hidden beaches to mountain retreats. Our guide takes you through the best spots for your next adventure.</p>
      </section>
      <section class="small-article1">
          <h3>Gear Up: Must-Have Gadgets</h3>
          <p>Check out the latest tech and gear to make your travels more exciting and comfortable.</p>
      </section>
      <section class="small-article2">
          <h3>Meet the Explorers</h3>
          <p>Get to know the modern adventurers who are pushing the boundaries of exploration.</p>
      </section>
      <section class="cover-image">
          <img src="https://cdn.freecodecamp.org/curriculum/labs/magazine-cover.png" alt="Adventure Cover Image">
      </section>
  </main>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 0;
  padding: 0;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f0f0f0;
}

.magazine-cover {
  display: grid;
  grid-template-areas: 
      "title title title"
      "feature-article feature-article cover-image"
      "small-article1 small-article2 cover-image";
  grid-template-columns: 1fr 1fr 1fr;
  grid-template-rows: auto 1fr 1fr;
  gap: 10px;
  width: 80%;
  background-color: #fff;
}

.title {
  grid-area: title;
  background-color: #1B1B32;
  color: white;
  padding: 10px;
  text-align: center;
}

.feature-article {
  grid-area: feature-article;
  background-color: #f9f9f9;
  padding: 10px;
}

.small-article1, .small-article2 {
  background-color: #e9e9e9;
  padding: 10px;
}

.cover-image {
  grid-area: cover-image;
  display: flex;
  justify-content: center;
  align-items: center;
  background-color: #ddd;
  padding: 10px;
}

.cover-image img {
  max-width: 100%;
  height: auto;
} 
```
