---
id: 5dc174fcf86c76b9248c6eb2
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在這個工作坊中，你將繼續使用標題、段落和列表等基本 HTML 元素，透過建置一個貓咪照片應用程式。 

開始工作坊，為你的頁面添加一個 `h1` 元素，內容為 `CatPhotoApp`。

# --hints--

程式碼中應該包含文字 `CatPhotoApp`。你可能想要檢查你的拼寫。

```js
assert.match(code, /catphotoapp/i);
```

你的 `h1` 元素應該有一個開啟標記。開啟標記具有這種語法：`<elementName>`。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 頁籤應該使用小寫。依慣例，所有 HTML 頁籤皆以小寫書寫。

```js
assert.notMatch(code, /<\/?H1>/);
```

你的 `h1` 元素應該有一個結束標記。結束標記具有此語法：`</elementName>`。

```js
assert.match(code, /<\/h1\>/);
```

你的 `h1` 元素的文字應該是 `CatPhotoApp`。你要麼遺漏了文字、有拼寫錯誤，要麼文字不在 `h1` 元素的開啟和關閉標籤之間。

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
