---
id: 5dc1798ff86c76b9248c6eb3
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `h1` 元素下方，添加帶有以下文本的 `h2` 元素：

`Cat Photos`

# --hints--

你的 `h1` 元素應該有一個開始標籤。 開始標籤的語法爲：`<elementName>`。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素應該有一個結束標籤。 結束標籤的語法爲：`</elementName>`。

```js
assert.match(code, /<\/h1\>/);
```

你應該只有一個 `h1` 元素。 刪除多餘的。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你的 `h1` 元素的文本應該是 “CatPhotoApp”。 你要麼遺漏了文本，要麼拼寫有誤。

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```

Your `h2` tags should be in lowercase. By convention, all HTML tags are written in lowercase.

```js
assert.notMatch(code, /<\/?H2>/);
```

你的 `h2` 元素應該有一個開始標籤。 開始標籤的語法爲：`<elementName>`。

```js
assert.exists(document.querySelector('h2'));
```

你的 `h2` 元素應該有一個結束標籤。 結束標籤在 `<` 字符之後有一個 `/`。

```js
assert.match(code, /<\/h2\>/);
```

你的 `h2` 元素的文本應該是 `Cat Photos`。 僅把文本 `Cat Photos` 放在 `h2` 的開始和結束標籤之間。

```js
assert.equal(document.querySelector('h2')?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat photos');
```

你的 `h2` 元素應該放在 `h1` 元素下面。 `h1` 元素更重要，必須位於 `h2` 元素之上。

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H1'), collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>

--fcc-editable-region--
  </body>
</html>
```
