---
id: 5dc1798ff86c76b9248c6eb3
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `h1` 元素下方，添加一個帶有此文字的 `h2` 元素：

`Cat Photos`

# --hints--

你的 `h1` 元素應該有一個開啟標記。開啟標記具有這種語法：`<elementName>`。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素應該有一個結束標記。結束標記具有此語法：`</elementName>`。

```js
assert.match(code, /<\/h1\>/);
```

你應該只保留一個 `h1` 元素。移除多餘的部分。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你的 `h1` 元素的文字應該是 'CatPhotoApp'。你要麼遺漏了文字，要麼有拼寫錯誤。

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```

你的 `h2` 頁籤應該使用小寫。依慣例，所有 HTML 頁籤皆以小寫書寫。

```js
assert.notMatch(code, /<\/?H2>/);
```

你的 `h2` 元素應該有一個開啟標籤。開啟標籤具有此語法：`<elementName>`。

```js
assert.exists(document.querySelector('h2'));
```

你的 `h2` 元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
assert.match(code, /<\/h2\>/);
```

你的 `h2` 元素的文字應該是 `Cat Photos`。只將文字 `Cat Photos` 放在開啟和關閉的 `h2` 頁籤之間。

```js
assert.equal(document.querySelector('h2')?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat photos');
```

你的 `h2` 元素應該位於 `h1` 元素之下。`h1` 元素具有較高的重要性，必須位於 `h2` 元素之上。

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H1'), collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
