---
id: 5dc24165f86c76b9248c6ebc
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

所有 `img` 元素都應該有一個 `alt` 屬性。`alt` 屬性的文字用於螢幕閱讀器以提升存取性，且在影像載入失敗時顯示。 

這是一個帶有 `alt` 屬性的 `img` 元素範例：

```html
<img src="cat.jpg" alt="A cat">
```

在 `img` 元素內，為你的 `alt` 屬性添加此文字：

`A cute orange cat lying on its back`

# --hints--

你的程式碼應該有一個 `img` 元素。你在先前的步驟中移除了 `img` 元素。

```js
assert.exists(document.querySelector('img'));
```

你的 `img` 元素沒有 `alt` 屬性。請檢查開啟標記名稱後是否有空格，和／或所有屬性名稱前是否有空格。

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

你的 `img` 元素的 `alt` 屬性值被設為除了 `A cute orange cat lying on its back` 以外的其他內容。請確保 `alt` 屬性值被引號包圍。

```js
const altText = document
  .querySelector('img')
  ?.alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert.match(altText, /A cute orange cat lying on its back\.?$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```
