---
id: 5dc2385ff86c76b9248c6eb7
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 有一些元素用來識別不同的內容區域。這些元素讓你的 HTML 更容易閱讀，並有助於搜尋引擎優化（SEO）和無障礙設計。

`main` 元素用於表述 HTML 文件主體的主要內容。`main` 元素內的內容應該是文件獨有的，且不應在文件的其他零件中重複出現。

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

透過在 `h1` 元素之前添加 `<main>` 開啟標籤，並在 `p` 元素之後添加 `</main>` 關閉標籤，來識別此頁面的主要區段。

# --hints--

你的 `main` 元素應該有一個開啟標記。開啟標記具有此語法：`<elementName>`。

```js
assert.exists(document.querySelector('main'));
```

你的 `main` 頁籤應該使用小寫。依慣例，所有 HTML 頁籤皆以小寫書寫。

```js
const mainTagMatches = code.matchAll(/<\/?main>/gi);
for (const match of mainTagMatches) {
  const tag = match[0]
  assert.strictEqual(
    tag,
    tag.toLowerCase()
  );
}
```

你的 `main` 元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
assert.match(code, /<\/main\>/);
```

你的 `main` 元素的開啟標籤應該在 `body` 元素的開啟標籤之下。你將它們的順序弄錯了。

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

你的 `main` 元素的開啟標籤應該在 `h1` 元素之上。你把它們的順序弄錯了。

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('MAIN'), collection.indexOf('H1'));
```

你的 `main` 元素的結束標籤應該在 `p` 元素的下方。你將它們的順序弄錯了。

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode?.contains(pNode));
assert.match(pNode?.textContent.toLowerCase(), /everyone loves cute cats online/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    
--fcc-editable-region--
  </body>
</html>
```
