---
id: 5dc24165f86c76b9248c6ebc
title: 步驟 9
challengeType: 0
dashedName: step-9
---

# --description--

所有的 `img` 元素都應該有一個 `alt` 屬性。 `alt` 屬性的文本用於屏幕閱讀器以改進可訪問性，以及在圖像加載失敗時顯示。

這是一個帶有 `alt` 屬性的 `img` 元素的例子：

```html
<img src="cat.jpg" alt="A cat">
```

在 `img` 元素中，添加一個 `alt` 屬性，文本爲：

`A cute orange cat lying on its back`

# --hints--

你的代碼應該有一個 `img` 元素。 你刪除了前一步驟中的 `img` 元素。

```js
assert.exists(document.querySelector('img'));
```

你的 `img` 元素應該有一個 `alt` 屬性。 檢查開始標籤的名稱後面是否有空格和/或所有屬性名稱之前是否有空格。

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

你的 `img` 元素的 `alt` 屬性值被設置成了其他內容，而非 `A cute orange cat lying on its back`。 確保將 `alt` 屬性的值放在引號中。

```js
const altText = document
  .querySelector('img')
  ?.alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert.match(altText, /A cute orange cat lying on its back\.?$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```
