---
id: 5dc23991f86c76b9248c6eb8
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

在上一步中，你將 `h1`、`h2`、註釋和 `p` 元素放在 `main` 元素中。 這叫作*嵌套*。 嵌套元素應該在其所嵌套元素的右邊再放兩個空格。 這種空格被稱爲縮進，用於使 HTML 更容易閱讀。

這是一個嵌套和縮進的示例：

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

下面代碼中的 `h1` 元素、`h2` 元素和註釋比 `main` 元素多縮進兩個空格。 使用鍵盤上的空格鍵在 `p` 元素前面添加兩個空格，使其正確縮進。

# --hints--

你的代碼應該有一個 `h2` 元素，文本爲 `Cat Photos`。 你可能不小心刪除了它，它缺少開始和結束標記，或者文本已經更改。

```js
assert.exists(document.querySelector('h2'));
assert.match(code, /<\/h2\>/);
assert.equal(document.querySelector('h2')?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(),'cat photos')
```

你不應添加示例中的 `ul` 或 `li` 元素。

```js
assert.notExists(document.querySelector('ul'));
assert.notExists(document.querySelector('li'));
```

你不應更改 `h2` 元素所在行的縮進。 它的開始標籤前應該有 6 個空格。 可以重置該步驟來恢復原始縮進。

```js
assert.match(code.toLowerCase(), /<\/h1\>\s*\n\s{6}<h2>/);
```

你的代碼應該有一個註釋。 你從之前的步驟中刪除了註釋。

```js
assert.match(code, /<!--.*-->/);
```

註釋的文本應該是 `TODO: Add link to cat photos`。 不要更改註釋的文本或空格。

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i);
```

你不應更改註釋元素所在行的縮進。 它的開始標籤前應該有 6 個空格。 可以重置該步驟來恢復原始縮進。

```js
assert.match(
  code
    .toLowerCase(),
    /<\/h2>\s*\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/
);
```

你的代碼應該有一個 `p` 元素。 你刪除了前一步驟中的 `p` 元素。

```js
assert.exists(document.querySelector('p'));
```

`p` 元素中的文本應該是 `Everyone loves cute cats online!`。不要改變 `p` 元素中的文本、空格和標點符號。

```js
assert.match(
  document
    .querySelector('p')
    ?.innerText?.trim().toLowerCase(),
    /^everyone\s+loves\s+cute\s+cats\s+online!$/
);
```

開始的 `p` 標籤應該有與 `h2` 和註釋元素一致的縮進。 它的開始標籤前應該有 6 個空格。

```js
assert.match(code.toLowerCase(), /-->\s*\n\s{6}<p>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
    <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
    </main>
  </body>
</html>
```
