---
id: 5dc23f9bf86c76b9248c6eba
title: 步驟 7
challengeType: 0
dashedName: step-7
---

# --description--

你可以通過使用 `img` 元素來爲你的網站添加圖像。 `img` 元素只有一個開始標籤，沒有結束標籤。 一個沒有結束標籤的元素被稱爲<dfn>空元素</dfn>。

在 `p` 元素下面添加一個 `img` 元素。 此時，在瀏覽器中還不會顯示任何圖像。

# --hints--

你的 `img` 元素應該有一個開始標籤。 開始標籤的語法爲：`<elementName>`。

```js
assert.exists(document.querySelector('img'));
```

你的 `img` 元素不應該有結束標籤 `</img>`。

```js
assert.notMatch(code, /<\/img\>/);
```

你應該只有一個 `img` 元素。 刪除任何多餘的。

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

你的 `img` 元素應該放在 `p` 元素下面。 你把它們的順序寫錯了。

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--

--fcc-editable-region--
    </main>
  </body>
</html>
```
