---
id: 5dc24614f86c76b9248c6ebd
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

你可以使用錨點（`a`）元素來聯結到另一個頁面。 

這裡有一個連結到 `https://www.freecodecamp.org` 的範例：

```html
<a href="https://www.freecodecamp.org"></a>
```

在段落後面添加一個錨點元素，該元素連結到 `https://freecatphotoapp.com`。此時，該連結不會顯示在預覽中。

# --hints--

你的錨點（`a`）元素應該有一個開啟標籤。開啟標籤具有此語法：`<elementName>`。

```js
assert.exists(document.querySelector('a'));
```

你的錨點（`a`）元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
assert.match(code, /<\/a\>/);
```

你的錨點（`a`）元素應該在 `p` 元素下面。你把它們的順序弄錯了。

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('A'));
```

你的錨點（`a`）元素沒有 `href` 屬性。請檢查開啟標籤名稱後是否有空格和／或所有屬性名稱前是否有空格。

```js
assert.isTrue(document.querySelector('a')?.hasAttribute('href'));
```

你的錨點（`a`）元素應該要連結到 `https://freecatphotoapp.com`。你要麼遺漏了 URL，要麼有拼寫錯誤。

```js
assert.equal(
  document.querySelector('a')?.getAttribute('href'),
    'https://freecatphotoapp.com'
);
```

雖然你已將錨點（`a`）元素的 `href` 屬性設為正確的聯結，但建議始終用引號將屬性值括起來。

```js
assert.notMatch(
  code, /\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
