---
id: 5dc23f9bf86c76b9248c6eba
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

你可以使用 `img` 元素為你的網站添加影像。`img` 元素只有開啟標籤，沒有關閉標籤。沒有關閉標籤的元素稱為 <dfn>void element</dfn>。

在 `p` 元素下方添加一個 `img` 元素。此時，瀏覽器中不會顯示任何影像。

# --hints--

你的 `img` 元素應該有一個開啟標籤。開啟標籤具有此語法：`<elementName>`。

```js
assert.exists(document.querySelector('img'));
```

你的 `img` 元素不應該有結束標籤 `</img>`。

```js
assert.notMatch(code, /<\/img\>/);
```

你應該只保留一個 `img` 元素。移除任何多餘的部分。

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

你的 `img` 元素應該在 `p` 元素下面。你把它們的順序弄錯了。

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
