---
id: 5dc24614f86c76b9248c6ebd
title: 步驟 10
challengeType: 0
dashedName: step-10
---

# --description--

你可以使用錨元素（`a`）鏈接到另一個頁面。

這是一個鏈接到 `https://www.freecodecamp.org` 的示例：

```html
<a href="https://www.freecodecamp.org"></a>
```

在段落之後添加一個鏈接到 `https://freecatphotoapp.com` 的錨元素。 此時，鏈接不會顯示在預覽中。

# --hints--

你的錨元素（`a`）應該有一個開始標籤。 開始標籤的語法爲：`<elementName>`。

```js
assert.exists(document.querySelector('a'));
```

你的錨元素（`a`）應該有一個結束標籤。 結束標籤在 `<` 字符之後有一個 `/`。

```js
assert.match(code, /<\/a\>/);
```

你的錨元素（`a`）應該在 `p` 元素下方。 你把它們的順序寫錯了。

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('A'));
```

你的錨元素（`a`）缺少 `href` 屬性。 檢查開始標籤的名稱後面是否有空格和/或所有屬性名稱之前是否有空格。

```js
assert.isTrue(document.querySelector('a')?.hasAttribute('href'));
```

你的錨元素（`a`）應該鏈接到 `https://freecatphotoapp.com`。 你要麼遺漏了 URL，要麼拼寫有誤。

```js
assert.equal(
  document.querySelector('a')?.getAttribute('href'),
    'https://freecatphotoapp.com'
);
```

雖然你將錨元素（`a`）的 `href` 屬性指向了正確的鏈接，但建議始終以引號來包裹屬性的值。

```js
assert.notMatch(
  code, /\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--

--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
