---
id: 5dfa3589eacea3f48c6300ae
title: 步驟 20
challengeType: 0
dashedName: step-20
---

# --description--

在第二個 `section` 元素中，添加一個帶有文本 `Cat Lists` 的新 `h2` 元素。

# --hints--

你的 `section` 元素應該有一個開始標籤。 開始標籤的語法爲：`<elementName>`。

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
assert.lengthOf(code.match(/<\/section>/g), 2);
```

你的 `h2` 元素應該有一個開始標籤。 開始標籤的語法爲：`<elementName>`。

```js
assert.lengthOf(document.querySelectorAll('h2'), 2);
```

你的 `h2` 元素應該有一個結束標籤。 結束標籤在 `<` 字符之後有一個 `/`。

```js
assert.lengthOf(code.match(/<\/h2\>/g), 2);
```

你的第二個 `h2` 元素應該在第二個 `section` 元素的結束標籤的上方。 它的位置不對。

```js
const secondSection = document.querySelectorAll('section')[1];
assert.equal(secondSection?.lastElementChild?.nodeName, 'H2');
```

第二個 `h2` 元素應該有文本 `Cat Lists`。 你要麼遺漏了文本，要麼拼寫有誤。

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    ?.lastElementChild?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>

      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
