---
id: 5dfa2407b521be39a3de7be1
title: 第 15 步
challengeType: 0
dashedName: step-15
---

# --description--

要在新分頁中開啟聯結，你可以在錨點（`a`）元素上使用 `target` 屬性。 

`target` 屬性指定要在哪裡開啟連結的文件。`target="_blank"` 會在新的分頁或視窗中開啟連結的文件。

以下是帶有 `target` 屬性的 `a` 元素的基本語法：

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

為你的錨點（`a`）元素的開啟標籤添加一個值為 `_blank` 的 `target` 屬性，使該連結在新分頁中開啟。

# --hints--

你的 `p` 元素應該包含一個巢狀的錨點（`a`）元素，且其文字為 `cat photos`。你可能刪除了它或有拼寫錯誤。

```js
const anchor = document.querySelectorAll('p > a')[1];
assert.equal(
    anchor?.innerText?.toLowerCase().trim().replace(/\s+/g, ' '), 'cat photos'
);
```

你的錨點（`a`）元素沒有 `target` 屬性。請檢查開啟標記名稱後是否有空格和／或所有屬性名稱前是否有空格。

```js
assert.isTrue(document.querySelectorAll('a')[1]?.hasAttribute('target'));
```

`target` 屬性的值應該是 `_blank`。你要麼遺漏了該值，要麼有拼寫錯誤。請記得屬性值應該用引號括起來。

```js
assert.equal(document.querySelectorAll('a')[1]?.getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
