---
id: 5dfb6250eacea3f48c6300b2
title: 第 24 步
challengeType: 0
dashedName: step-24
---

# --description--

在無序列表之後，添加一個新的影像，其 `src` 屬性值設定為：

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

以及它的 `alt` 屬性值為：

`A slice of lasagna on a plate.`

# --hints--

在關閉的 `</ul>` 標籤後面應該有一個 `img` 元素。

```js
assert.equal(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.nodeName, 'IMG');
```

新的影像沒有 `alt` 屬性。請檢查開啟標記名稱後是否有空格和／或所有屬性名稱前是否有空格。

```js
assert.isTrue(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.hasAttribute('alt'));
```

新的影像應該有一個 `alt` 值為 `A slice of lasagna on a plate.`。請確保 `alt` 屬性的值被引號包圍。

```js
assert.match(
  document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.getAttribute('alt')?.replace(/\s+/g, ' '), /^A slice of lasagna on a plate\.?$/i 
);
```

新的影像沒有 `src` 屬性。請檢查開啟標記名稱後是否有空格和／或所有屬性名稱前是否有空格。

```js
assert.isTrue(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.hasAttribute('src'));
```

新的影像應該有一個 `src` 值為 `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`。請確保 `src` 屬性的值被引號包圍。

```js
assert.equal(
  document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.getAttribute('src'),
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

雖然你已將新影像的 `src` 設定為正確的 URL，但建議始終用引號將屬性的值括起來。

```js
assert.notMatch(code, /\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
