---
id: 5dfb6250eacea3f48c6300b2
title: 步驟 24
challengeType: 0
dashedName: step-24
---

# --description--

在無序列表之後，添加一個新圖像，它具有 `src` 屬性，值爲：

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

而它的 `alt` 屬性值爲：

`A slice of lasagna on a plate.`

# --hints--

在結束 `</ul>` 標籤之後應該有一個 `img` 元素。

```js
assert.equal(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.nodeName, 'IMG');
```

新圖像沒有 `alt` 屬性。 檢查開始標記的名稱後面是否有空格和/或所有屬性名稱之前是否有空格。

```js
assert.isTrue(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.hasAttribute('alt'));
```

新圖像的 `alt` 值應爲 `A slice of lasagna on a plate.`。確保將 `alt` 屬性的值放在引號中。

```js
assert.match(
  document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.getAttribute('alt')?.replace(/\s+/g, ' '), /^A slice of lasagna on a plate\.?$/i 
);
```

新圖像沒有 `src` 屬性。 檢查開始標籤的名稱後面是否有空格和/或所有屬性名稱之前是否有空格。

```js
assert.isTrue(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.hasAttribute('src'));
```

新圖像的 `src` 值應爲 `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`。 確保將 `src` 屬性的值放在引號中。

```js
assert.equal(
  document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.getAttribute('src'),
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

雖然你已將新圖像的 `src` 設置爲正確的 URL，但建議始終將屬性的值用引號括起來。

```js
assert.notMatch(code, /\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
