---
id: 5ef9b03c81a63668521804d0
title: 步驟 27
challengeType: 0
dashedName: 步驟-27
---

# --description--

要強調一個特定的單詞或短語，你可以使用 `em` 元素。

通過將 `figcaption` 元素中的單詞 `love` 包裹在強調 `em` 元素中來強調它。

# --hints--

你的強調 `em` 元素應該有一個開始標籤。 開始標籤的語法爲：`<elementName>`。

```js
assert.exists(document.querySelector('em'));
```

你的強調 `em` 元素應該有一個結束標籤。 結束標籤在 `<` 字符之後有一個 `/`。

```js
assert.match(code, /<\/em\>/);
```

你要麼刪除了 `figcaption` 元素，要麼它缺少開始或結束標籤。

```js
assert.exists(document.querySelector('figcaption'));
assert.match(code, /<\/figcaption\>/);
```

你的強調 `em` 元素應包裹文本 `love`。 你要麼省略了文本，要麼有拼寫錯誤。

```js
assert.equal(
  document.querySelector('figcaption > em')?.innerText?.trim().toLowerCase(), 'love'
);
```

`figcaption` 的文本應該是 `Cats love lasagna`。 檢查是否有拼寫錯誤，以及 `em` 元素的開始和結束標籤周圍是否有必要的空格。

```js
assert.match(
  document
    .querySelector('figcaption')
    ?.innerText?.trim().replace(/\s+/gi, ' '),
    /^cats love lasagna\.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```
