---
id: 5ef9b03c81a63668521804e7
title: 步驟 35
challengeType: 0
dashedName: step-35
---

# --description--

`footer` 元素用於定義文件或區段的頁腳。頁腳通常包含文件作者的資訊、版權資料、使用條款的連結、聯絡資訊等。

在 `main` 元素之後，為你的文件添加 `footer` 元素。

# --hints--

你已刪除 `main` 元素，或它缺少開啟標籤或關閉標籤。

```js
assert.exists(document.querySelector('main'));
assert.match(code, /<\/main>/);
```

你的 `footer` 元素應該有一個開啟標記。開啟標記具有以下語法：`<elementName>`。

```js
assert.exists(document.querySelector('footer'));
```

你的 `footer` 元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
assert.match(code, /<\/footer\>/);
```

你的 `footer` 元素應該位於關閉的 `main` 元素標籤下方。你把它放在其他地方了。

```js
assert.equal(document.querySelector('main')?.nextElementSibling.nodeName, 'FOOTER');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
