---
id: 5ef9b03c81a63668521804e8
title: 步驟 36
challengeType: 0
dashedName: step-36
---

# --description--

在 `footer` 元素中嵌套一個 `p` 元素，其文本爲 `No Copyright - freeCodeCamp.org`。

# --hints--

你要麼刪除了 `footer` 元素，要麼它缺少開始標籤或結束標籤。

```js
assert.exists(document.querySelector('footer'));
assert.match(code, /<\/footer>/);
```

你的 `footer` 元素應該包含一個 `p` 元素。 確保爲 `p` 元素添加一個開始標籤和結束標籤。

```js
assert.exists(document.querySelector('footer > p'));
```

你的 `footer` 元素應該有一個結束標籤。 結束標籤在 `<` 字符之後有一個 `/`。

```js
const pElemClosingTags = code.match(/<\/p\>/g);
assert.exists(pElemClosingTags);
assert.lengthOf(pElemClosingTags, 3);
```

你的 `p` 元素的文本應該是 `No Copyright - freeCodeCamp.org`。 你要麼遺漏了文本，要麼拼寫有誤，要麼沒有把文本放在 `p` 元素的開始和結束標籤之間。

```js
const extraSpacesRemoved = document.querySelectorAll('footer > p')[0]?.innerText?.trim().replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /^No Copyright - freeCodeCamp\.org$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
--fcc-editable-region--
    <footer>

    </footer>
--fcc-editable-region--
  </body>
</html>
```
