---
id: 5ef9b03c81a63668521804ea
title: 第 38 步
challengeType: 0
dashedName: step-38
---

# --description--

請注意，到目前為止你添加到頁面上的所有內容都在 `body` 元素內。所有應該呈現到頁面的頁面內容元素都放在 `body` 元素內。然而，其他重要的資訊則放在 `head` 元素內。

`head` 元素用於包含關於文件的詮釋資料，例如其標題、樣式表的連結和腳本。詮釋資料是關於頁面的資訊，但不會直接顯示在頁面上。

在 `body` 元素上方為你的文件添加一個 `head` 元素。

# --hints--

你已刪除 `body` 元素，或它缺少開啟標籤或關閉標籤。

```js
assert.exists(document.querySelector('body'));
assert.match(code, /<\/body>/);
```

你的 `head` 元素應該有一個開啟標記。開啟標記具有以下語法：`<elementName>`。

```js
assert.match(code, /\<head\>/);
```

你的 `head` 元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
assert.match(code, /\<\/head\>/);
```

你的 `head` 元素應該位於開啟的 `body` 元素標籤之上。你把它放在其他地方了。

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<\/head\>\<body\>/);
```

# --seed--

## --seed-contents--

```html
<html>
--fcc-editable-region--
  
--fcc-editable-region--
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
