---
id: 5ef9b03c81a63668521804ea
title: 步驟 38
challengeType: 0
dashedName: step-38
---

# --description--

注意，到目前爲止添加到頁面中的所有內容都位於`body`元素中。 所有應該呈現到頁面的頁面內容元素都放在 `body` 元素中。 但是，其他重要信息在 `head` 元素中。

`head` 元素用於包含關於文檔的元數據，如它的標題、樣式錶鏈接和腳本。 元數據是不直接顯示在頁面上的關於頁面的信息。

在 `body` 元素上方添加一個 `head` 元素。

# --hints--

你要麼刪除了 `body` 元素，要麼它缺少開始標籤或結束標籤。

```js
assert.exists(document.querySelector('body'));
assert.match(code, /<\/body>/);
```

你的 `head` 元素應該有一個開始標籤。 開始標籤的語法爲：`<elementName>`。

```js
assert.match(code, /\<head\>/);
```

你的 `head` 元素應該有一個結束標籤。 結束標籤在 `<` 字符之後有一個 `/`。

```js
assert.match(code, /\<\/head\>/);
```

你的 `head` 元素應位於 `body` 元素開始標籤的上方。 你把它放在別的地方了。

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<\/head\>\<body\>/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<html>

  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
