---
id: 5ef9b03c81a63668521804ec
title: 步驟 40
challengeType: 0
dashedName: step-40
---

# --description--

請注意，整個頁面的內容都巢狀在一個 `html` 元素中。`html` 元素是 HTML 頁面的根元素，並包裹頁面上的所有內容。

你也可以透過為 `html` 元素添加 `lang` 屬性來指定你的頁面語言。 

為你的開啟 `html` 頁籤添加值為 `en` 的 `lang` 屬性，以指定頁面的語言為英文。

# --hints--

你已刪除 `html` 元素，或它缺少開啟標籤或關閉標籤。

```js
assert.match(code, /\<html.*?\>/);
assert.match(code, /\<\/html\>/);
```

你的 `html` 元素應該有一個 `lang` 屬性，且其值為 `en`。你可能遺漏了該屬性／值，或有拼寫錯誤。

```js
const extraSpacesRemoved = code.replace(/\s+/g, " ");
assert.match(extraSpacesRemoved, /\<html\s+lang\s*\=\s*("?|'?)en\1\s*\>/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<html>
--fcc-editable-region--
  <head>
    <title>CatPhotoApp</title>
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
