---
id: 5ef9b03c81a63668521804ee
title: 步驟 41
challengeType: 0
dashedName: step-41
---

# --description--

所有頁面都應以 `<!DOCTYPE html>` 開頭。這個特殊字串稱為 <dfn>宣告、宣告式</dfn>，可確保瀏覽器嘗試符合整個產業的規格。

`<!DOCTYPE html>` 告訴瀏覽器該文件是 HTML5 文件，這是 HTML 的最新版本。

將此宣告式添加到程式碼的第一行。

# --hints--

你的程式碼應該以宣告式 `<!DOCTYPE html>` 開頭。你可能遺漏了宣告式、打錯字，或它不是程式碼的第一行。

```js
assert.match(code, /<\s*!doctype\s+html\s*>/i);
```

你的 `<!DOCTYPE html>` 必須位於文件頂端。

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /^<\!doctypehtml><html/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
<html lang="en">
  <head>
    <title>CatPhotoApp</title>
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
