---
id: 5ef9b03c81a63668521804e9
title: 步驟 37
challengeType: 0
dashedName: step-37
---

# --description--

通過將現有的 `freeCodeCamp.org` 文本包含在錨元素（`a`）中，使其成爲鏈接。 `href` 屬性應該被設置爲 `https://www.freecodecamp.org`。

# --hints--

你的錨元素（`a`）應該嵌套在 `footer` 中的 `p`  元素中。 確保爲錨元素（`a`）添加一個開始和結束標籤。

```js
assert.isNotEmpty(document.querySelectorAll('footer > p > a'));
```

你的錨元素（`a`）應該有一個結束標籤。 結束標籤在 `<` 字符之後有一個 `/`。

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert.exists(aElemClosingTags);
assert.lengthOf(aElemClosingTags, 4);
```

你的錨元素（`a`）應該有一個 `href` 屬性，其值爲 `https://www.freecodecamp.org`。 你要麼遺漏了屬性/值，要麼拼寫有誤。

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(nestedAnchor?.getAttribute('href').toLowerCase(), 'https://www.freecodecamp.org');
```

鏈接的文本應該是 `freeCodeCamp.org`。 你要麼遺漏了文本，要麼拼寫有誤。

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(
  nestedAnchor?.innerText?.trim().toLowerCase(),
    'freecodecamp.org'
);
```

嵌套錨元素（`a`）後，瀏覽器中唯一可見的 `p` 元素內容應該是 `No Copyright - freeCodeCamp.org`。 仔細檢查 `p` 和嵌套錨元素的文本、空格或標點符號。

```js
const pText = document.querySelectorAll('footer > p')[0]?.innerText?.trim().toLowerCase().replace(/\s+/g, ' ');
assert.match(pText, /^no copyright - freecodecamp\.org$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```
