---
id: 5f07be6ef7412fbad0c5626b
title: 步驟 18
challengeType: 0
dashedName: step-18
---

# --description--

在添加任何新內容之前，你應該使用 `section` 元素將貓咪照片內容與未來內容分開。

`section` 元素用於定義文件中的區段，例如章節、標頭、頁尾或文件中的其他任何區段。它是一個語意元素，有助於 SEO 和無障礙設計。

```html
<section>
  <h2>Section Title</h2>
  <p>Section content...</p>
</section>
```

將你的 `h2` 元素、兩個 `p` 元素和錨點（`a`）元素巢狀在一個 `section` 元素中。

# --hints--

你的 `section` 元素應該有一個開啟標記。開啟標記具有以下語法：`<elementName>`。

```js
assert.exists(document.querySelector('section'));
```

你的 `section` 元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
assert.match(code, /<\/section\s*>/i);
```

整個 `section` 元素應該位於 `main` 元素的開啟和關閉標籤之間。

```js
assert.equal(document.querySelector('section')?.parentNode.nodeName, 'MAIN');
```

現有的 `h2`、`p` 元素和錨點（`a`）元素應該位於 `section` 元素的開啟與關閉標籤之間。

```js
const sectionEl = document.querySelector('section');
assert.isNotNull(sectionEl);
const childrenOfSection = [...sectionEl.childNodes];

const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert.lengthOf(foundElements, 4)
```

`h1` 元素不應該巢狀在 `section` 元素中。

```js
const sectionEl = document.querySelector('section');
assert.isNotNull(sectionEl);
const childrenOfSection = [...sectionEl.childNodes];

const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      
      <h2>Cat Photos</h2>
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
