---
id: 5f07c98cdb9413cbd4b16750
title: 第 19 步
challengeType: 0
dashedName: step-19
---

# --description--

是時候添加一個新區段。將第二個 `section` 元素添加到現有的 `section` 元素下方。

# --hints--

你的 `section` 元素應該有一個開啟標籤。開啟標籤具有此語法：`<elementName>`。

```js
assert.isAtLeast(document.querySelectorAll('section').length, 2);
```

你應該只添加一個開啟的 `section` 頁籤。請移除任何多餘的。

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
```

你的 `section` 元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
assert.isAtLeast(code.match(/<\/section>/g)?.length, 2);
```

你應該只添加一個結束的 `section` 頁籤。請移除任何多餘的。

```js
assert.lengthOf(code.match(/<\/section>/g), 2);
```

第二個 `section` 元素不應該巢狀在第一個 `section` 元素中。

```js
const el = document.querySelector('main > section');
assert.isNotNull(el);
const childrenOf1stSection = [
  ...el.children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.isEmpty(foundElems);
```

兩個 `section` 元素都應該位於 `main` 元素的開啟與關閉標籤之間。

```js
const el = document.querySelector('main');
assert.isNotNull(el);
const childrenOfMain = [...el.children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.lengthOf(foundElems, 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
