---
id: 671141feba228a35cefba82d
title: 步驟14
challengeType: 0
dashedName: step-14
---

# --description--

將現有文本 `cute cats` 變成一個錨元素，鏈接到：

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`

# --hints--

在第一個 `p` 元素中應該有一個新的錨元素。

```js
assert.exists(document.querySelector('p')?.querySelector('a'));
```

錨文本應爲 `cute cats`。

```js
assert.strictEqual(document.querySelector('a')?.innerText?.trim(), "cute cats");
```

新的錨元素的 `href` 屬性應爲 `https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`。

```js
assert.strictEqual(document.querySelector('a')?.href, "https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg");
```

在你的代碼中應該只有一個短語 `cute cats` 。

```js
assert.strictEqual(code.indexOf("cute cats"),code.lastIndexOf("cute cats")); 
```

`p` 元素的文本仍應是 `Everyone loves cute cats online!`。

```js
assert.strictEqual(document.querySelector('p')?.innerText?.trim(), "Everyone loves cute cats online!");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
