---
id: 671b6e873249bb35c9debfcf
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

在錨點元素前添加字、單字 `See more `，在錨點元素後添加字、單字 ` in our gallery`。

# --hints--

你應該使用現有的錨點元素，勿創建新的。

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

錨點元素內的文字不應該改變，必須是 `cat photos`。

```js
assert.strictEqual(document.querySelector('a')?.innerText?.replace(/\s+/g, ' ').trim(), "cat photos");
```

你應該在你的程式碼中有 `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery`。

```js
assert.match(code, /See more\s+<a href="https:\/\/freecatphotoapp\.com">cat photos<\/a>\s+in our gallery/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
