---
id: 587d825a367417b2b2512c87
title: Erstelle eine doppelt-verkettete Liste
challengeType: 1
forumTopicId: 301626
dashedName: create-a-doubly-linked-list
---

# --description--

All of the linked lists we've created so far are singly linked lists. Here, we'll create a <dfn>doubly linked list</dfn>. As the name implies, nodes in a doubly linked list have references to the next and previous node in the list.

Dies erlaubt es uns die Liste in beide Richtungen zu traversieren, erfordert jedoch auch mehr Speicher, da jeder Knoten eine zusätzliche Referenz auf den vorherigen Knoten speichern muss.

# --instructions--

Wir haben dir ein `Node` Objekt zur Verfügung gestellt und unsere `DoublyLinkedList` initialisiert. Lass uns nun zwei Methoden `add` und `remove` zu unserer doppelt-verketteten Liste hinzufügen. Die Methode `add` soll ein übergebenes Element zur Liste hinzufügen, während die Methode `remove` alle Vorkommnisse des übergebenen Elements aus der Liste entfernen soll.

Achte beim Schreiben dieser Methoden darauf, alle möglichen Randfälle zu behandeln, z. B. Löschungen für das erste oder letzte Element. Zudem sollte beim Versuch ein Element aus einer leeren Liste zu entfernen, `null` zurückgegeben werden.

# --hints--

Die `DoublyLinkedList` Datenstruktur sollte vorhanden sein.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return typeof test == 'object';
  })()
);
```

Die `DoublyLinkedList` sollte eine Methode namens `add` besitzen.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.add == undefined) {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

Die `DoublyLinkedList` sollte eine Methode namens `remove` besitzen.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.remove == undefined) {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

Beim Entfernen eines Elements aus einer leeren Liste sollte `null` zurückgegeben werden.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return test.remove(100) == null;
  })()
);
```

Die `add`-Methode sollte Elemente zur Liste hinzufügen.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(5);
    test.add(6);
    test.add(723);
    return test.print().join('') == '56723';
  })()
);
```

Jeder Knoten sollte eine Referenz zum vorherigen Knoten besitzen.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(50);
    test.add(68);
    test.add(73);
    return test.printReverse().join('') == '736850';
  })()
);
```

Das erste Element der Liste sollte entfernbar sein.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(25);
    test.add(35);
    test.add(60);
    test.remove(25);
    return test.print().join('') == '3560';
  })()
);
```

Das letzte Element der Liste sollte entfernbar sein.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(25);
    test.add(35);
    test.add(60);
    test.remove(60);
    return test.print().join('') == '2535';
  })()
);
```

# --seed--

## --after-user-code--

```js
DoublyLinkedList.prototype = Object.assign(
  DoublyLinkedList.prototype,
  {

  print() {
    if (this.head == null) {
      return null;
    } else {
      var result = new Array();
      var node = this.head;
      while (node.next != null) {
        result.push(node.data);
        node = node.next;
      };
      result.push(node.data);
      return result;
    };
  },
  printReverse() {
    if (this.tail == null) {
      return null;
    } else {
      var result = new Array();
      var node = this.tail;
      while (node.prev != null) {
        result.push(node.data);
        node = node.prev;
      };
      result.push(node.data);
      return result;
    };
  }
});
```

## --seed-contents--

```js
var Node = function(data, prev) {
  this.data = data;
  this.prev = prev;
  this.next = null;
};
var DoublyLinkedList = function() {
  this.head = null;
  this.tail = null;
  // Only change code below this line

  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
