---
id: 66da326c02141df538f29ba5
title: 創建一個結賬頁面
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

**目標：**實現以下用戶需求並通過所有測試以完成實驗。

**用戶需求：**

1. 你應該有一個 `h1` 元素，文本爲 `Checkout` 。
1. `h1` 元素之後應緊跟兩個 `section` 元素。
1. 在第一部分中，應包含一個 `h2` 元素，文本爲 `Your Cart`。
1. 應在第一部分提供項目圖片，並配以適當的備用文字。 如果你願意，可以使用這張圖片：`https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. 在第二部分中，應包含一個 `h2` 元素，文本爲 `Payment Information`。
1. 在第二部分中，應有一個 `form` 元素。
1. 在表單中，應包含一個輸入元素，其 `id` 和 `name` 屬性均爲 `card-name`，並且有與之關聯的 `label`。
1. 在表單中，應包含一個輸入元素，其 `id` 和 `name` 屬性均爲 `card-number`，並且有與之關聯的 `label`。
1. 你至少應有兩個輸入元素是必填項，並且應將 `aria-required` 屬性設置爲 `true`。

# --hints--

你應該有一個 `h1` 元素，文本爲 `Checkout` 。

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

你的頁面上應該只有一個 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

在 `h1` 元素之後，你應至少包含兩個 `section` 元素。

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

在第一部分中，應包含一個 `h2` 元素，文本爲 `Your Cart`。

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

在第一部分中，應該至少有一張帶有替代文本的圖片。

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

在第二部分中，應包含一個 `h2` 元素，文本爲 `Payment Information`。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

在第二個 `section` 中，應包含一個 `form` 元素。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

在表單中，應包含一個 `input` 元素，其 `id` 和 `name` 爲 `card-name`。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

在表單中，應包含一個 `input` 元素，其 `id` 和 `name` 爲 `card-number`。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

所有非 `submit` `type` 的 `input` 元素，都應有一個與之關聯的 `label` 元素。

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => assert.exists(document.querySelector(`label[for="${input.id}"]`)));
```

你至少應有兩個 `input` 元素是必填項，並且應將 `aria-required` 屬性設置爲 `true`。

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required && input.getAttribute('aria-required') === 'true');
assert.isAtLeast(requiredInputs?.length, 2)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name</label>
                    <input type="text" id="card-name" name="card-name" required aria-required="true">
                </div>
                <div>
                    <label for="card-number">Card Number</label>
                    <input type="text" id="card-number" name="card-number" required aria-required="true">
                </div>
                <div>
                    <label for="expiry-date">Expiry Date</label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required aria-required="true">
                </div>
                <div>
                    <label for="cvv">CVV</label>
                    <input type="text" id="cvv" name="cvv" required aria-required="true" aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```
