---
id: 66da326c02141df538f29ba5
title: 創建一個結賬頁面
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

**目標：**實現以下用戶需求並通過所有測試以完成實驗。

**用戶需求：**

1. 你應該有一個 `h1` 元素，文本爲 `Checkout` 。
1. `h1` 元素之後應緊跟兩個 `section` 元素。
1. 在第一部分中，應包含一個 `h2` 元素，文本爲 `Your Cart`。
1. 應在第一部分提供項目圖片，並配以適當的備用文字。 如果你願意，可以使用這張圖片：`https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. 在第二部分中，應包含一個 `h2` 元素，文本爲 `Payment Information`。
1. 在第二部分中，應有一個 `form` 元素。
1. 在表單中，應包含一個輸入元素，其 `id` 和 `name` 屬性均爲 `card-name`，`type` 爲 `text`，並且有與之關聯的 `label`。
1. 在表單中，應包含一個輸入元素，其 `id` 和 `name` 屬性均爲 `card-number`，`type` 爲 `text`，並且有與之關聯的 `label`。
1. 你至少應包含兩個帶有  `required` 屬性的 `input`元素。
1. 你在每個必填輸入的 `label` 元素內，都應該有一個 `span` 元素，文本爲  `*`，`aria-hidden` 設置爲 `true`，以便必填字段能通過視覺方式標明。
1. 你應在卡號輸入框的正後方放置一個 `p` 元素，其中包含說明所需卡號格式的幫助文本。 `p` 應該有一個值爲 `card-number-help` 的 `id`，並被卡號輸入框使用 `aria-describedby` 進行引用。

# --hints--

你應該有一個 `h1` 元素，文本爲 `Checkout` 。

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

你的頁面上應該只有一個 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

在 `h1` 元素之後，你應至少包含兩個 `section` 元素。

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

在第一部分中，應包含一個 `h2` 元素，文本爲 `Your Cart`。

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

在第一部分中，應該至少有一張帶有替代文本的圖片。

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

在第二部分中，應包含一個 `h2` 元素，文本爲 `Payment Information`。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

在第二個 `section` 中，應包含一個 `form` 元素。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

在表單中，應包含一個 `input` 元素，其 `id` 和 `name` 爲 `card-name`。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

你的卡片名稱輸入框應該有一個值爲 `text` 的 `type` 屬性。

```js
const cardNameInput = document.getElementById('card-name');
assert.exists(cardNameInput);
assert.equal(cardNameInput.getAttribute('type'), 'text');
```

在表單中，應包含一個 `input` 元素，其 `id` 和 `name` 爲 `card-number`。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

你的卡號輸入框應該有一個值爲 `text` 的 `type` 屬性。

```js
const cardNumberInput = document.getElementById('card-number');
assert.exists(cardNumberInput);
assert.equal(cardNumberInput.getAttribute('type'), 'text');
```

所有非 `submit` `type` 的 `input` 元素，都應有一個與之關聯的 `label` 元素。

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
});
```

你至少應包含兩個帶有 `required` 屬性的 `input` 元素。

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required);
assert.isAtLeast(requiredInputs?.length, 2)
```

對於每個必填輸入項，應在 `label` 元素內包含一個帶有文本 `*` 的 `span` 元素。

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.textContent.trim(), '*');
});
```

你的 `span` 元素應將 `aria-hidden` 屬性設置爲 `true`。

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.getAttribute('aria-hidden'), 'true');
});
```

在卡號輸入框的正下方，應放置一個具有 `id="card-number-help"` 屬性的 `p` 元素。

```js
const cardNumberInput = document.getElementById('card-number');
const cardNumberHelp = document.getElementById('card-number-help');

assert.exists(cardNumberHelp);
const nextElement = cardNumberInput.parentElement?.tagName === 'LABEL'
  ? cardNumberInput.parentElement.nextElementSibling
  : cardNumberInput.nextElementSibling;
assert.equal(nextElement, cardNumberHelp);
```

你的卡號幫助文本不應爲空。

```js
const cardNumberHelp = document.getElementById('card-number-help');
assert.isString(cardNumberHelp?.textContent);
assert.isAbove(cardNumberHelp?.textContent.trim().length, 0);
```

你的卡號輸入框應設置 `aria-describedby` 屬性爲 `card-number-help`。

```js
const cardNumberInput = document.getElementById('card-number');
assert.equal(cardNumberInput.getAttribute('aria-describedby'), 'card-number-help');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-name" name="card-name" required>
                </div>
                <div>
                    <label for="card-number">Card Number <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-number" name="card-number" aria-describedby="card-number-help" required>
                    <p id="card-number-help">Please enter your 16-digit card number without spaces or dashes.</p>
                </div>
                <div>
                    <label for="expiry-date">Expiry Date <span aria-hidden="true">*</span></label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required>
                </div>
                <div>
                    <label for="cvv">CVV <span aria-hidden="true">*</span></label>
                    <input type="text" id="cvv" name="cvv" required aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```
