---
id: 66da326c02141df538f29ba5
title: 建置結帳頁面
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `h1` 元素，裡面包含文字 `Checkout`。
1. 你應該在 `h1` 元素之後緊接著有兩個 `section` 元素。
1. 你應該在第一個區段內有一個 `h2` 元素，其文字為 `Your Cart`。
1. 你應該在第一部分有一張帶有適當替代文字的項目影像。你可以使用這張影像：`https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. 你應該在第二個區段中有一個 `h2` 元素，其 `text` 為 `Payment Information`。
1. 你應該在第二個區段中有一個 `form` 元素。
1. 你應該在你的表單中有一個 `id` 和 `name` 為 `card-name`，且 `type` 為 `text` 的輸入欄位，以及一個與之相關的 `label`。
1. 你應該在你的表單中有一個 `id` 和 `name` 為 `card-number`，且 `type` 為 `text` 的輸入欄位，以及一個與之相關的 `label`。
1. 你應該至少有兩個帶有 `required` 屬性的 `input` 元素。
1. 你應該在每個必填輸入的 `label` 元素內包含一個 `span` 元素，其 `文字` 為 `*` 且 `aria-hidden` 設為 `true`，以便在視覺上顯示必填欄。
1. 你應該有一個 `p` 元素，內含說明所需卡號格式的說明文字，並放置在卡號輸入欄位的正後方。該 `p` 應具有 `id` 為 `card-number-help`，並由卡號輸入欄位使用 `aria-describedby` 取用。

# --hints--

你應該有一個 `h1` 元素，裡面包含文字 `Checkout`。

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

你的頁面上應該只有一個 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你應該在 `h1` 元素之後至少有兩個 `section` 元素。

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

你應該在第一個區段內有一個 `h2` 元素，其文字為 `Your Cart`。

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

你的第一個章節中應該至少有一個帶有替代文字的影像。

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

你應該在第二個區段中有一個 `h2` 元素，其 `text` 為 `Payment Information`。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

你應該在第二個 `section` 元素中有一個 `form` 元素。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

你應該在你的表單中有一個 `input` 元素，其 `id` 和 `name` 為 `card-name`。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

你的卡片名稱輸入應該具有 `type` 為 `text`。

```js
const cardNameInput = document.getElementById('card-name');
assert.exists(cardNameInput);
assert.equal(cardNameInput.getAttribute('type'), 'text');
```

你應該在你的表單中有一個 `input` 元素，其 `id` 和 `name` 為 `card-number`。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

你的卡號輸入應該具有 `type` 為 `text`。

```js
const cardNumberInput = document.getElementById('card-number');
assert.exists(cardNumberInput);
assert.equal(cardNumberInput.getAttribute('type'), 'text');
```

所有不是 `submit` `type` 的 `input` 元素都應該有相關的 `label` 元素。

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
});
```

你應該至少有兩個帶有 `required` 屬性的 `input` 元素。

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required);
assert.isAtLeast(requiredInputs?.length, 2)
```

你應該在每個必填輸入的 `label` 元素內包含一個帶有文字 `*` 的 `span` 元素。

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.textContent.trim(), '*');
});
```

你的 `span` 元素應該將 `aria-hidden` 設定為 `true`。

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.getAttribute('aria-hidden'), 'true');
});
```

你應該在卡號輸入欄位後立即放置一個 `id` 為 `card-number-help` 的 `p` 元素。

```js
const cardNumberInput = document.getElementById('card-number');
const cardNumberHelp = document.getElementById('card-number-help');

assert.exists(cardNumberHelp);
const nextElement = cardNumberInput.parentElement?.tagName === 'LABEL'
  ? cardNumberInput.parentElement.nextElementSibling
  : cardNumberInput.nextElementSibling;
assert.equal(nextElement, cardNumberHelp);
```

你的卡號說明文字不得為空。

```js
const cardNumberHelp = document.getElementById('card-number-help');
assert.isString(cardNumberHelp?.textContent);
assert.isAbove(cardNumberHelp?.textContent.trim().length, 0);
```

你的卡號輸入欄位應該將 `aria-describedby` 設定為 `card-number-help`。

```js
const cardNumberInput = document.getElementById('card-number');
assert.equal(cardNumberInput.getAttribute('aria-describedby'), 'card-number-help');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-name" name="card-name" required>
                </div>
                <div>
                    <label for="card-number">Card Number <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-number" name="card-number" aria-describedby="card-number-help" required>
                    <p id="card-number-help">Please enter your 16-digit card number without spaces or dashes.</p>
                </div>
                <div>
                    <label for="expiry-date">Expiry Date <span aria-hidden="true">*</span></label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required>
                </div>
                <div>
                    <label for="cvv">CVV <span aria-hidden="true">*</span></label>
                    <input type="text" id="cvv" name="cvv" required aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```
