---
id: 5f332a88dc25a0fd25c7687a
title: Hatua ya 6
challengeType: 0
dashedName: step-6
---

# --description--

The name of the cafe is `CAMPER CAFE`. So, add an `h1` element within your `main` element. Give it the name of the cafe in capitalized letters to make it stand out.

# --hints--

Unapaswa kuwa na tagi ya kufungua ya `<h1>`.

```js
assert.match(code, /<h1>/i);
```

Unapaswa kuwa na tagi ya kufunga ya `</h1>`.

```js
assert.match(code, /<\/h1>/i);
```

Hupaswi kubadilisha kipengele chako cha `main` kilichopo.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

Kipengee chako cha `h1` kinapaswa kuwekwa kwenye kipengee chako cha `main`.

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi `CAMPER CAFE`.

```js
assert.match(code, /<h1>CAMPER CAFE<\/h1>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

