---
id: 61fd77f7ad2aeb5ae34d07d6
title: 步骤14
challengeType: 0
dashedName: step-14
---

# --description--

在你的第二个 `tr` 元素中，添加一个带有文本 `Checking Our primary transactional account.` 的 `th` 元素。将该文本中除 `Checking ` 之外的部分包裹在一个 `class` 属性设置为 `description` 的 `span` 元素中。

接着，添加三个 `td` 元素，并添加以下文本（按顺序）：`$54`、`$56`、`$53`。 给第三个 `td` 元素添加一个 `class` 属性，设置为 `current`。

# --hints--

你的第二个 `tr` 应该有一个 `th` 元素。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.isNotNull(tableRow?.querySelector('th'));
```

你的 `th` 元素应包含文本 `Checking Our primary transactional account.`

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.strictEqual(
  tableRow?.querySelector('th')?.innerText.trim(),
  'Checking Our primary transactional account.'
);
```

你应该将文本 `Our primary transactional account.` 包裹在 `span` 元素中。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.strictEqual(
  tableRow?.querySelector('th > span')?.textContent.trim(),
  'Our primary transactional account.'
);
```

你的 `span` 应该有一个 `class` 属性，设置为 `description`。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.isTrue(
  tableRow?.querySelector('th > span')?.classList?.contains('description')
);
```

你应该有三个 `td` 元素。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.lengthOf(tableRow?.querySelectorAll('td'), 3);
```

你的第一个 `td` 元素应该包含文本 `$54`。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[0]?.textContent.trim(), '$54');
```

你的第二个 `td` 元素应该包含文本 `$56`。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[1]?.textContent.trim(), '$56');
```

你的第三个 `td` 元素应该包含文本 `$53`。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[2]?.textContent.trim(), '$53');
```

你的第三个 `td` 元素应该将 `class` 设置为 `current`。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.isTrue(
  tableRow?.querySelectorAll('td')?.[2]?.classList?.contains('current')
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="data"></tr>
              <tr class="total"></tr>
            </tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
