---
id: 61fd67a656743144844941cb
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

屏幕阅读器根据文档流阅读 HTML 元素。 我们最终希望资产负债表有一个 “Balance Sheet” 的标题和一个 “AcmeWidgetCorp” 的子标题。 但是，如果屏幕阅读器按此顺序阅读则没有意义。

将现有的 `span` 元素的 `class` 属性设置为 `flex`，并在其内部添加两个 `span` 元素。 给第一个添加文本 `AcmeWidgetCorp`。 给第二个添加文本 `Balance Sheet`。 你将使用 CSS 来颠倒页面上文字的顺序，但 HTML 的顺序对于屏幕阅读器来说更合理。

# --hints--

你现有的 `span` 元素应将 `class` 属性设置为 `flex`。

```js
assert.isTrue(
  document.querySelector('h1')?.children?.[0]?.classList?.contains('flex')
);
```

你现有的 `span` 元素中应有两个新的 `span` 元素。

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.localName,
  'span'
);
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.localName,
  'span'
);
```

你的新 `span` 元素不应有 `class` 属性。

```js
assert.isEmpty(document.querySelector('.flex')?.children?.[0]?.classList);
assert.isEmpty(document.querySelector('.flex')?.children?.[1]?.classList);
```

你的第一个新 `span` 元素应包含文本 `AcmeWidgetCorp`。

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.textContent.trim(),
  'AcmeWidgetCorp'
);
```

你的第二个新 `span` 元素应包含文本 `Balance Sheet`。

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.textContent.trim(),
  'Balance Sheet'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
