---
id: 61fd5a93fd62bb35968adeab
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

通过在 `head` 元素中为样式表添加 `link` 元素，开始你的项目。

# --hints--

你的代码中应该有一个 `link` 元素。

```js
assert.match(code, /<link/);
```

你的 `link` 元素应该在你的 `head` 元素中。

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

你的 `link` 元素应该有一个 `rel` 属性，值为 `stylesheet`。

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/);
```

你的 `link` 元素应该有一个 `href` 属性，值为 `styles.css`。

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
