---
id: 61fd6ab779390f49148773bb
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

在你的 `h1` 元素下方，创建一个 `div` 元素。 给它一个 `id` 属性，设置为 `years`。 你希望对屏幕阅读器隐藏此特定元素，因此将 `aria-hidden` 属性设置为 `true`。

# --hints--

应该在 `h1` 元素之后创建一个新的 `div` 元素。

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling?.localName,
  'div'
);
```

你的新 `div` 元素应将 `id` 属性设置为 `years`。

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('id'), 'years');
```

新的新 `div` 元素应将 `aria-hidden` 属性设置为 `true`。

```js
assert.strictEqual(
  document.querySelector('div')?.getAttribute('aria-hidden'),
  'true'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
