---
id: 61fd6cc9475a784b7776233e
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

在你现有的 `div` 元素下方，添加一个新的 `div` 元素，并将其 `class` 属性设置为 `table-wrap`。 这将是你的表格的容器。

在其中添加三个 `table` 元素。 你将使用 CSS 将这些样式设置为单个表格，但使用单独的表格将有助于屏幕阅读器理解文档流。

# --hints--

你应创建一个新的 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

你的新 `div` 元素应该有一个设置为 `table-wrap` 的 `class` 。

```js
assert.strictEqual(document.querySelector('.table-wrap')?.localName, 'div');
```

你的 `.table-wrap` 元素应该在现有的 `div` 之后。

```js
assert.isTrue(
  document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap')
);
```

你的 `.table-wrap` 元素应该包含三个 `table` 元素。

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert.lengthOf(children, 3);
children.forEach(child => assert.strictEqual(child?.localName, 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
