---
id: 61fd6b7c83dbf54a08cf0498
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

在你的 `div` 元素中，添加三个 `span` 元素。 为它们每个设置一个 `class` 属性，并将其设置为 `year`，并添加以下文本（按顺序）：`2019`、`2020` 和 `2021`。

# --hints--

你的 `div` 元素应该有三个 `span` 元素。

```js
assert.lengthOf(document.querySelector('div')?.children, 3);
```

每个 `span` 元素应该有一个设置为 `year` 的 `class` 属性。

```js
const spans = [...document.querySelector('div')?.children];
spans.forEach(span => assert.isTrue(span?.classList?.contains('year')));
```

你的第一个 `span` 应该包含文本 `2019`。

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[0]?.textContent.trim(),
  '2019'
);
```

你的第二个 `span` 应该包含文本 `2020`。

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[1]?.textContent.trim(),
  '2020'
);
```

你的第三个 `span` 应该包含文本 `2021`。

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[2]?.textContent.trim(),
  '2021'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--
        <div id="years" aria-hidden="true">
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
