---
id: 61fd9ad665a4a282c8106be3
title: 步骤 26
challengeType: 0
dashedName: step-26
---

# --description--

在你的 `thead` 中，创建一个 `tr` 元素。 在其中，添加一个 `td` 和三个 `th` 元素。 在每个 `th` 元素中，添加一个 `span` 元素，并将 `class` 设置为 `sr-only`，并按顺序添加以下文本：`2019`、`2020` 和 `2021`。

# --hints--

你的 `thead` 元素应该有一个 `tr` 元素。

```js
const table = document.querySelectorAll('table')?.[2];
assert.strictEqual(
  table?.querySelector('thead')?.children?.[0]?.localName,
  'tr'
);
```

`tr` 元素应该有一个 `td` 元素作为第一个子元素。

```js
const table = document.querySelectorAll('table')?.[2];
assert.strictEqual(table?.querySelector('tr')?.children?.[0]?.localName, 'td');
```

`tr` 元素应该具有三个 `th` 元素，在 `td` 元素之后。

```js
const table = document.querySelectorAll('table')?.[2];
const trElement = table?.querySelector('tr');
assert.strictEqual(trElement?.children?.[1]?.localName, 'th');
assert.strictEqual(trElement?.children?.[2]?.localName, 'th');
assert.strictEqual(trElement?.children?.[3]?.localName, 'th');
```

你的每个 `th` 元素都应该有一个 `span` 元素。

```js
const ths = [
  ...document.querySelectorAll('table')?.[2]?.querySelectorAll('th')
];
ths?.forEach(th => {
  assert.lengthOf(th?.children, 1);
  assert.strictEqual(th?.children?.[0]?.localName, 'span');
});
```

你的每个新的 `span` 元素都应将 `class` 属性设置为 `sr-only`。

```js
const ths = [
  ...document.querySelectorAll('table')?.[2]?.querySelectorAll('th')
];
ths?.forEach(th => {
  assert.isTrue(th?.children?.[0]?.classList?.contains('sr-only'));
});
```

你的第一个 `span` 元素应包含文本 `2019`。

```js
const table = document.querySelectorAll('table')?.[2];
const thead = table.querySelectorAll('th')?.[0];
assert.strictEqual(thead?.children?.[0]?.textContent.trim(), '2019');
```

你的第二个 `span` 元素应包含文本 `2020`。

```js
const table = document.querySelectorAll('table')?.[2];
const thead = table.querySelectorAll('th')?.[1];
assert.strictEqual(thead?.children?.[0]?.textContent.trim(), '2020');
```

你的第三个 `span` 元素应包含文本 `2021`。

```js
const table = document.querySelectorAll('table')?.[2];
const thead = table.querySelectorAll('th')?.[2];
assert.strictEqual(thead?.children?.[0]?.textContent.trim(), '2021');
```

你的 `td` 元素应该为空。

```js
const table = document.querySelectorAll('table')?.[2];
const tableData = table?.querySelectorAll('td')?.[0];
assert.strictEqual(tableData?.textContent.trim(), '');
assert.isTrue(tableData?.children?.length === 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>Savings <span class="description">Funds set aside for emergencies.</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
              <tr class="total">
                <th>Total <span class="sr-only">Assets</span></th>
                <td>$579</td>
                <td>$736</td>
                <td class="current">$809</td>
              </tr>
            </tbody>
          </table>
          <table>
            <caption>
              Liabilities
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only">2019</span></th>
                <th><span class="sr-only">2020</span></th>
                <th><span class="sr-only">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Loans <span class="description">The outstanding balance on our startup loan.</span></th>
                <td>$500</td>
                <td>$250</td>
                <td class="current">$0</td>
              </tr>
              <tr class="data">
                <th>Expenses <span class="description">Annual anticipated expenses, such as payroll.</span></th>
                <td>$200</td>
                <td>$300</td>
                <td class="current">$400</td>
              </tr>
              <tr class="data">
                <th>Credit <span class="description">The outstanding balance on our credit card.</span></th>
                <td>$50</td>
                <td>$50</td>
                <td class="current">$75</td>
              </tr>
              <tr class="total">
                <th>Total <span class="sr-only">Liabilities</span></th>
                <td>$750</td>
                <td>$600</td>
                <td class="current">$475</td>
              </tr>
            </tbody>
          </table>
--fcc-editable-region--
          <table>
            <caption>Net Worth</caption>
            <thead>
            </thead>
            <tbody>
            </tbody>
          </table>
--fcc-editable-region--
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
