---
id: 68507cadbf36aa089a84ce2d
title: 루언 알고리즘 구현하기
challengeType: 27
dashedName: implement-the-luhn-algorithm
---

# --description--

루언 알고리즘은 "modulus 10" 또는 "mod 10" 알고리즘으로도 알려져 있으며, 신용카드 번호 같은 다양한 식별 번호를 검증하는 데 사용하는 간단한 체크섬 공식입니다. 루언 알고리즘으로 번호를 검증하는 단계는 다음과 같습니다:

- 오른쪽 끝에서 시작하여, 가장 오른쪽 숫자(체크 숫자)를 제외하고, 한 자리씩 건너뛰며 각 숫자의 값을 두 배로 만드세요.
- 두 배로 만든 숫자가 `9`보다 크면, 각 자리 숫자를 더해 한 자리 숫자로 만드세요. 또는 결과에서 `9`을 빼도 됩니다.
- 체크 숫자를 포함한 모든 숫자의 합을 구하세요.
- 모든 숫자의 합이 `10`의 배수이면 번호가 유효한 것이고, 그렇지 않으면 유효하지 않습니다.

예를 들어 `453914881`라는 번호가 있다고 합시다. 루언 알고리즘으로 검증하는 단계는 다음과 같습니다:

 ```md 
 계좌 번호           4   5   3   9   1   4   8   8   1 
 한 자리씩 건너뛰며 두 배  4  10   3  18   1   8   8  16   1 
 두 자리 숫자 합산       4   1   3   9   1   8   8   7   1 
 ``` 

- 그런 다음 모든 숫자를 더하면 `4 + 1 + 3 + 9 + 1 + 8 + 8 + 7 + 1 = 42`가 됩니다.
- `42`가 `10`의 배수가 아니므로 번호는 유효하지 않습니다.

이 실습에서는 루언 알고리즘을 사용해 신용카드 번호 검증기를 만드세요.

**목표:** 아래 사용자 스토리를 완성하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. `verify_card_number`라는 이름의 함수를 정의해야 하며, 이 함수는 숫자 문자열(카드 번호를 나타냄)을 받아 루언 알고리즘에 따라 유효한지 검증해야 합니다.

2. `verify_card_number` 함수 내에서:

   - 전달된 카드 번호에 포함될 수 있는 대시(-)나 공백을 처리해야 합니다.
   - 카드 번호가 유효하면 `VALID!`를 반환하고, 그렇지 않으면 `INVALID!`를 반환하세요.

프로젝트를 완료하면 입력에 따라 다음 메시지가 표시됩니다:

| 카드 번호            | 메시지   |
| ------------------- | -------- |
| 453914889           | VALID!   |
| 4111-1111-1111-1111 | VALID!   |
| 1234 5678 9012 3456 | INVALID! |

# --hints--

`verify_card_number`라는 이름의 함수가 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('verify_card_number')
    `)
})
```

`verify_card_number('453914889')`는 `VALID!`를 반환해야 합니다.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914889') == 'VALID!'
    `)
})
```

`verify_card_number('4111-1111-1111-1111')`는 `VALID!`를 반환해야 합니다.

```js
({
    test: () => runPython(`
    assert verify_card_number('4111-1111-1111-1111') == 'VALID!'
    `)
})
```

`verify_card_number('453914881')`는 `INVALID!`를 반환해야 합니다.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914881') == 'INVALID!'
    `)
})
```

`verify_card_number('1234 5678 9012 3456')`는 `INVALID!`를 반환해야 합니다.

```js
({
    test: () => runPython(`
    assert verify_card_number('1234 5678 9012 3456') == 'INVALID!'
    `)
})
```

`verify_card_number`는 유효한 신용카드 번호로 호출되었을 때 `VALID!`를 반환해야 합니다.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6467') == 'VALID!'
    assert verify_card_number('6011 1111 1111 1117') == 'VALID!'
    assert verify_card_number('3782-822463-10005') == 'VALID!'
    assert verify_card_number('5555 5555 5555 4444') == 'VALID!'
    assert verify_card_number('3056 9309 0259 04') == 'VALID!'
    assert verify_card_number('6011000990139424173') == 'VALID!'
    `)
})
```

`verify_card_number`는 유효하지 않은 신용카드 번호로 호출되었을 때 `INVALID!`를 반환해야 합니다.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6466') == 'INVALID!'
    assert verify_card_number('6011 1111 1111 1116') == 'INVALID!'
    assert verify_card_number('3782 822463 10004') == 'INVALID!'
    assert verify_card_number('5555-5555-5555-4440') == 'INVALID!'
    assert verify_card_number('3056 9309 0259 05') == 'INVALID!'
    assert verify_card_number('6011000990139424175') == 'INVALID!'
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def verify_card_number(card_number):
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    sum_of_odd_digits = 0
    card_number_reversed = translated_card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            number = (number // 10) + (number % 10)
        sum_of_even_digits += number

    total = sum_of_odd_digits + sum_of_even_digits
    if total % 10 == 0:
        return 'VALID!'
    else:
        return 'INVALID!'

# Example usage
verify_card_number('4111-1111-4555-1142')
```
