---
id: 646ce8bb4b35544d501c7184
title: 20단계
challengeType: 0
dashedName: step-20
---

# --description--

`.cat-ears` 요소 안에 각각 `div`와 `cat-left-ear` 클래스를 가진 두 개의 `cat-right-ear` 요소를 만드세요.

# --hints--

기존 `div` 요소와 `cat-ears` 클래스는 변경하지 마세요.

```js
assert(document.querySelectorAll('div.cat-ears').length === 1);
```

`div` 요소 안에 두 개의 `.cat-ears` 요소를 만들어야 합니다.

```js
assert(document.querySelectorAll('.cat-ears div').length === 2);
```

첫 번째 `div` 요소에는 `cat-left-ear` 클래스를 지정하세요.

```js
assert(document.querySelectorAll('.cat-ears div')[0]?.classList.contains('cat-left-ear'));
```

두 번째 `div` 요소에는 `cat-right-ear` 클래스를 지정하세요.

```js
assert(document.querySelectorAll('.cat-ears div')[1]?.classList.contains('cat-right-ear'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        --fcc-editable-region--
        <div class="cat-ears">
          
        </div>
        --fcc-editable-region--
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
