---
id: 646c47867800472a4ed5d2ea
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

기본 HTML 구조부터 시작하세요. `DOCTYPE` 참조 `html`와 `html` 속성이 `lang`로 설정된 `en` 요소를 추가하세요. 또한 `head` 요소 안에 `body`와 `html` 요소를 추가하세요.

# --hints--

`DOCTYPE`의 `html` 선언이 있어야 합니다.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

HTML 맨 앞에 `DOCTYPE` 선언이 있어야 합니다.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`html` 속성이 `lang`인 `en` 여는 태그가 있어야 합니다.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

닫는 `html` 태그가 있어야 합니다.

```js
assert(code.match(/<\/html>/i));
```

여는 `head` 태그가 있어야 합니다.

```js
assert(code.match(/<head>/i));
```

닫는 `head` 태그가 있어야 합니다.

```js
assert(code.match(/<\/head>/i));
```

여는 `body` 태그가 있어야 합니다.

```js
assert(code.match(/<body>/i));
```

닫는 `body` 태그가 있어야 합니다.

```js
assert(code.match(/<\/body>/i));
```

`body` 요소 다음에 `head` 요소가 와야 합니다.

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

`head`와 `body` 요소가 `html` 요소 안에 있어야 합니다.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
