---
id: 646c48df8674cf2b91020ecb
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`head` 요소 안에 `meta` 태그를 `charset` 속성이 `utf-8`인 상태로 추가하세요. 또한 `title` 요소를 텍스트 `fCC Cat Painting`와 함께 추가하세요.

# --hints--

하나의 `meta` 요소를 추가하세요.

```js
assert(document.querySelectorAll('meta').length === 1);
```

`meta` 요소에 `charset` 속성을 `utf-8`로 설정하세요.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

하나의 `title` 요소를 추가하세요.

```js
assert(document.querySelectorAll('title').length === 1);
```

`title` 요소에 텍스트 `fCC Cat Painting`를 넣으세요. 철자와 대소문자가 정확해야 합니다.

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
