---
id: 646c48df8674cf2b91020ecc
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

`link` 요소 안에 `head` 요소를 추가하세요. 그 `link` 요소에 `rel` 속성을 `stylesheet`로, `href`를 `./styles.css`로 설정하세요.

# --hints--

`link` 요소가 있어야 합니다.

```js
assert.match(code, /<link/)
```

`link` 요소에는 `rel`와 `href` 속성이 각각 `stylesheet`과 `./styles.css`로 설정되어야 합니다.

```js
assert.match(code, /<link\s+(rel=("|')stylesheet\2\s+href=("|')\.\/styles\.css\3|href=("|')\.\/styles\.css\4\s+rel=("|')stylesheet\5)\s*\/?>/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```

```css

```
