---
id: 646c5ace05e4be4211407935
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

`.cat-head` 요소를 보려면 `#5e5e5e`을 `85%`에, `#45454f`을 `100%`에 지정한 선형 그라디언트 배경을 주세요.

이 두 색상 차이를 눈치채지 못할 수도 있지만, 분명히 존재합니다.

# --hints--

`.cat-head` 선택자에 `background` 속성이 있어야 합니다.

```js
assert.match(code, /background:/)
```

`background` 속성은 `linear-gradient` 함수를 사용해야 합니다.

```js
assert.match(code, /background:\s*linear-gradient\(/)
```

`linear-gradient` 함수는 첫 번째 색상을 `#5e5e5e`에 `85%`로 설정해야 합니다.

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const firstGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[1];
assert.equal(firstGradient,'(94,94,94)85%,');
```

`linear-gradient` 함수는 두 번째 색상을 `#45454f`에 `100%`로 설정해야 합니다.

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const secondGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[2];
assert.equal(secondGradient,'(69,69,79)100%)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;

}
--fcc-editable-region--
```
