---
id: 646ce9d790d2a44de5f99e04
title: 21단계
challengeType: 0
dashedName: step-21
---

# --description--

두 개의 `div` 요소를 만드세요. 첫 번째는 `.cat-left-ear` 클래스를 가진 `cat-left-inner-ear` 요소 안에, 두 번째는 `.cat-right-ear` 클래스를 가진 `cat-right-inner-ear` 요소 안에 있어야 합니다.

# --hints--

`div` 클래스를 가진 기존 `cat-left-ear` 요소는 변경하지 마세요.

```js
assert(document.querySelectorAll('div.cat-left-ear').length === 1);
```

`div` 클래스를 가진 기존 `cat-right-ear` 요소는 변경하지 마세요.

```js
assert(document.querySelectorAll('div.cat-right-ear').length === 1);
```

`div` 요소 안에 `.cat-left-ear` 요소가 하나 있어야 합니다.

```js
assert(document.querySelectorAll('.cat-left-ear div').length === 1);
```

`div` 요소 안에 `.cat-right-ear` 요소가 하나 있어야 합니다.

```js
assert(document.querySelectorAll('.cat-right-ear div').length === 1);
```

`div` 안에 새로 만든 `.cat-left-ear` 요소는 `cat-left-inner-ear` 클래스를 가져야 합니다.

```js
assert(document.querySelectorAll('.cat-left-ear div')[0]?.classList.contains('cat-left-inner-ear'));
```

`div` 안에 새로 만든 `.cat-right-ear` 요소는 `cat-right-inner-ear` 클래스를 가져야 합니다.

```js
assert(document.querySelectorAll('.cat-right-ear div')[0]?.classList.contains('cat-right-inner-ear'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          --fcc-editable-region--
          <div class="cat-left-ear">

          </div>
          <div class="cat-right-ear">
            
          </div>
          --fcc-editable-region--
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
