---
id: 646c5d7057c45f432fcdd46c
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

클래스 선택자를 사용해 `.cat-head` 요소에 너비 `205px`와 높이 `180px`를 지정하세요. 또한, 테두리 `1px solid #000`와 `border-radius`의 `46%`도 지정하세요.

# --hints--

`.cat-head` 선택자가 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

`.cat-head` 선택자에 `width`가 `205px`로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.width === '205px');
```

`.cat-head` 선택자에 `height`가 `180px`로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.height === '180px')
```

`.cat-head` 선택자에 `border`가 `1px solid #000`로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.border === '1px solid rgb(0, 0, 0)')
```

`.cat-head` 선택자에 `border-radius`가 `46%`로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document)?.getStyle(".cat-head")?.borderRadius === '46%')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class='cat-head'></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--

--fcc-editable-region--
```
