---
id: 646dd556d524bc61c0139bd6
title: 35단계
challengeType: 0
dashedName: step-35
---

# --description--

대부분 고양이는 같은 귀의 바깥과 안쪽 부분 색이 다릅니다. 똑같이 구현할 수 있습니다. 그래서 조금 전에 오른쪽과 왼쪽 안쪽 귀에 `div` 요소를 정의했죠.

클래스 선택자를 사용해 `.cat-left-inner-ear` 요소에 좌우 테두리를 각각 `20px solid transparent`로 주세요. 그리고 아래쪽 테두리는 `40px solid #3b3b4f`로 주세요.

# --hints--

`.cat-left-inner-ear` 선택자가 있어야 합니다.

```js 
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-inner-ear'))
```

`.cat-left-inner-ear` 선택자에 `border-left` 속성을 `20px solid transparent`로 설정하세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-inner-ear')?.borderLeft === '20px solid transparent')
```

`.cat-left-inner-ear` 선택자에 `border-right` 속성을 `20px solid transparent`로 설정하세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-inner-ear')?.borderRight === '20px solid transparent')
```

`.cat-left-inner-ear` 선택자에 `border-bottom` 속성을 `40px solid #3b3b4f`로 설정하세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-inner-ear')?.borderBottom === '40px solid rgb(59, 59, 79)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  z-index: 1;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  position: absolute;
  top: -26px;
  left: 163px;
  z-index: 1;
  transform: rotate(45deg);
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

--fcc-editable-region--

--fcc-editable-region--
```
