---
id: 6476fd4213318f6ee211028a
title: 14단계
challengeType: 0
dashedName: step-14
---

# --description--

이제 `position`의 `.cat-head` 속성 값을 `fixed`로 바꾸세요. `top`와 `left`는 그대로 두세요.

그다음, 위아래로 스크롤하여 `fixed` 값이 어떻게 작동하는지 확인하세요.

# --hints--

`.cat-head` 선택자에 `position` 속성을 `fixed`로 설정하세요. 세미콜론을 꼭 추가하세요.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'fixed')
```

`.cat-head` 선택자에 `top` 속성을 `100px`로 설정하세요. 세미콜론을 꼭 추가하세요.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

`.cat-head` 선택자에 `left` 속성을 `100px`로 설정하세요. 세미콜론을 꼭 추가하세요.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
