---
id: 6477062778c85972eb648030
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

이제 고양이 머리를 가운데에 배치하세요.

`.cat-head` 요소에 `position` 속성을 `absolute`로 설정하세요. `0`, `right`, `left`, `top` 속성에는 `bottom` 값을 설정한 다음, 모든 면에 `margin` 속성을 `auto`로 설정하세요. 이것이 CSS 위치 지정으로 요소를 수직과 수평 모두 가운데에 배치하는 한 가지 방법입니다.

# --hints--

`.cat-head` 선택자에 `position` 속성을 `absolute`로 설정하세요. 세미콜론을 꼭 추가하세요.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

`.cat-head` 선택자에 `top` 속성을 `0`로 설정하세요. 세미콜론을 꼭 추가하세요.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

`.cat-head` 선택자에 `left` 속성을 `0`로 설정하세요. 세미콜론을 꼭 추가하세요.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '0px')
```

`.cat-head` 선택자에 `right` 속성을 `0`로 설정하세요. 세미콜론을 꼭 추가하세요.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.right === '0px')
```

`.cat-head` 선택자에 `bottom` 속성을 `0`로 설정하세요. 세미콜론을 꼭 추가하세요.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.bottom === '0px')
```

`.cat-head` 선택자에 `margin` 속성을 `auto`로 설정하세요. 세미콜론을 꼭 추가하세요.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.margin === 'auto')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: sticky;
  top: 0;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 400px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
