---
id: 647d855ac12cd436059acd39
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

범용 선택자를 사용해 CSS에 `box-sizing: border-box;`를 추가하세요. 이렇게 하면 요소가 지정된 너비와 높이에 패딩과 테두리를 포함하게 됩니다.

# --hints--

범용 선택자(`*`)가 있어야 합니다.

```js
// assert.match(code, /\*\s*{\s*.*\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*'))
```

범용 선택자에 `box-sizing` 속성이 `border-box`로 설정되어야 합니다.

```js
// assert.match(code, /\*\s*{\s*box-sizing:\s*border-box;?\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
