---
id: 587d78af367417b2b2512b03
title: Criar um formulário de pesquisa
challengeType: 14
dashedName: build-a-survey-form
demoType: onClick
---

# --description--

Atenda às histórias de usuário abaixo e faça todos os testes passarem para completar o laboratório.

**Histórias de usuário:**

1. Você deve ter um título de página em um elemento `h1` com o `id` `title`.
1. Você deve ter uma breve explicação em um elemento `p` com o `id` `description`.
1. Você deve ter um elemento `form` com um `id` de `survey-form`.
1. Dentro do elemento de formulário, deve haver um campo de entrada requerido `input` para inserir seu nome, que tenha um `id` de `name` e um `type` de `text`.
1. Dentro do elemento de formulário, deve haver um campo de entrada requerido `input` para inserir seu e-mail, que tenha um `id` de `email`.
1. Se for informado um e-mail que não esteja formatado corretamente, um erro de validação HTML5 deve ser mostrado.
1. Dentro do elemento de formulário, deve haver um campo de entrada `input` para inserir um número que tenha um `id` de `number`.
1. A entrada de número não deve aceitar algo que não seja números, impedindo que você os digite ou mostrando um erro de validação do HTML5 (dependendo do seu navegador).
1. Se forem inseridos números fora do intervalo do campo de entrada do número, intervalo esse definido pelos atributos `min` e `max`, um erro de validação de HTML5 deve ser mostrado.
1. Para os campos de entrada nome, e-mail e número dentro do formulário, deve haver elementos `label` correspondentes, que descrevam o propósito de cada campo com os seguintes ids: `name-label`, `email-label` e `number-label`.
1. Para os campos de entrada de nome, e-mail e número, deve haver um texto placeholder que forneça uma descrição ou instruções para cada campo.
1. Dentro do elemento do formulário, você deve ter um elemento de menu suspenso `select` com um `id` `dropdown` e pelo menos duas opções para escolher.
1. Dentro do elemento do formulário, você pode selecionar uma opção de um grupo de pelo menos dois botões de opção (radio) que serão agrupados usando o atributo `name`.
1. Dentro do elemento de formulário, deve ser possível selecionar vários campos de uma série de caixas de seleção (checkboxes). Cada um desses campos deve ter um atributo `value`.
1. Dentro do elemento de formulário, deve haver uma `textarea` para comentários adicionais.
1. Dentro do elemento de formulário, deve haver um botão com o `id` `submit` para enviar todas as entradas.

# --hints--

Você deve ter um elemento `h1` com um `id` de `title`.

```js
const el = document.getElementById('title');
assert(!!el && el.tagName === 'H1');
```

O elemento `#title` não deve estar vazio.

```js
const el = document.getElementById('title');
assert(!!el && el.innerText.length > 0);
```

Você deve ter um elemento `p` com um `id` de `description`.

```js
const el = document.getElementById('description');
assert(!!el && el.tagName === 'P');
```

O elemento `#description` não deve estar vazio.

```js
const el = document.getElementById('description');
assert(!!el && el.innerText.length > 0);
```

Você deve ter um elemento `form` com um `id` de `survey-form`.

```js
const el = document.getElementById('survey-form');
assert(!!el && el.tagName === 'FORM');
```

Você deve ter um elemento `input` com um `id` de `name`.

```js
const el = document.getElementById('name');
assert(!!el && el.tagName === 'INPUT');
```

O seu `#name` deve ter um `type` de `text`.

```js
const el = document.getElementById('name');
assert(!!el && el.type === 'text');
```

O elemento `#name` deve exigir a entrada.

```js
const el = document.getElementById('name');
assert(!!el && el.required);
```

O elemento `#name` deve ser um descendente de `#survey-form`.

```js
const el = document.querySelector('#survey-form #name');
assert(!!el);
```

Você deve ter um elemento `input` com um `id` de `email`.

```js
const el = document.getElementById('email');
assert(!!el && el.tagName === 'INPUT');
```

O elemento `#email` deve ter um `type` de `email`.

```js
const el = document.getElementById('email');
assert(!!el && el.type === 'email');
```

O elemento `#email` deve exigir a entrada.

```js
const el = document.getElementById('email');
assert(!!el && el.required);
```

O elemento `#email` deve ser um descendente de `#survey-form`.

```js
const el = document.querySelector('#survey-form #email');
assert(!!el);
```

Você deve ter um elemento `input` com um `id` de `number`.

```js
const el = document.getElementById('number');
assert(!!el && el.tagName === 'INPUT');
```

O elemento `#number` deve ser um descendente de `#survey-form`.

```js
const el = document.querySelector('#survey-form #number');
assert(!!el);
```

O elemento `#number` deve ter um `type` de `number`.

```js
const el = document.getElementById('number');
assert(!!el && el.type === 'number');
```

O elemento `#number` deve ter um atributo `min` com um valor numérico.

```js
const el = document.getElementById('number');
assert(!!el && el.min && isFinite(el.min));
```

O elemento `#number` deve ter um atributo `max` com um valor numérico.

```js
const el = document.getElementById('number');
assert(!!el && el.max && isFinite(el.max));
```

Você deve ter um elemento `label` com um `id` de `name-label`.

```js
const el = document.getElementById('name-label');
assert(!!el && el.tagName === 'LABEL');
```

Você deve ter um elemento `label` com um `id` de `email-label`.

```js
const el = document.getElementById('email-label');
assert(!!el && el.tagName === 'LABEL');
```

Você deve ter um elemento `label` com um `id` de `number-label`.

```js
const el = document.getElementById('number-label');
assert(!!el && el.tagName === 'LABEL');
```

`#name-label` deve conter um texto que descreva a entrada.

```js
const el = document.getElementById('name-label');
assert(!!el && el.innerText.length > 0);
```

`#email-label` deve conter um texto que descreva a entrada.

```js
const el = document.getElementById('email-label');
assert(!!el && el.innerText.length > 0);
```

`#number-label` deve conter um texto que descreva a entrada.

```js
const el = document.getElementById('number-label');
assert(!!el && el.innerText.length > 0);
```

O elemento `#name-label` deve ser um descendente de `#survey-form`.

```js
const el = document.querySelector('#survey-form #name-label');
assert(!!el);
```

O elemento `#email-label` deve ser um descendente de `#survey-form`.

```js
const el = document.querySelector('#survey-form #email-label');
assert(!!el);
```

O elemento `#number-label` deve ser um descendente de `#survey-form`.

```js
const el = document.querySelector('#survey-form #number-label');
assert(!!el);
```

O elemento `#name` deve ter o atributo `placeholder` e um valor.

```js
const el = document.getElementById('name');
assert(!!el && !!el.placeholder && el.placeholder.length > 0);
```

O elemento `#email` deve ter o atributo `placeholder` e um valor.

```js
const el = document.getElementById('email');
assert(!!el && !!el.placeholder && el.placeholder.length > 0);
```

O elemento `#number` deve ter o atributo `placeholder` e um valor.

```js
const el = document.getElementById('number');
assert(!!el && !!el.placeholder && el.placeholder.length > 0);
```

Você deve ter um campo `select` com um `id` de `dropdown`.

```js
const el = document.getElementById('dropdown');
assert(!!el && el.tagName === 'SELECT');
```

O elemento `#dropdown` deve ter pelo menos dois elementos selecionáveis (não desativados) `option`.

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert(els.length >= 2);
```

O elemento `#dropdown` deve estar aninhado dentro de `#survey-form`.

```js
const el = document.querySelector('#survey-form #dropdown');
assert(!!el);
```

Você deve ter pelo menos dois elementos `input` com `type` `radio` (botões de opção).

```js
const els = document.querySelectorAll('input[type="radio"]');
assert(els.length >= 2);
```

Você deve ter pelo menos dois botões de opção aninhados dentro de `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert(els.length >= 2);
```

Todos os seus botões de opção devem ter um atributo `value` e um valor.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll('input[type="radio"][value=""], input[type="radio"]:not([value])');
assert(els1.length > 0 && els2.length === 0);
```

Todos os seus botões de opção devem ter um atributo `name` e um valor.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll('input[type="radio"][name=""], input[type="radio"]:not([name])');
assert(els1.length > 0 && els2.length === 0);
```

Todos os grupos de botões de opção devem ter pelo menos 2 botões de opção.

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  })
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert(false);
})

assert(groupKeys.length > 0);
```

Você deve ter pelo menos dois elementos `input` com `type` `checkbox` (caixas de seleção) aninhados dentro de `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert(els.length >= 2);
```

Todas as caixas de seleção dentro de `#survey-form` devem ter um atributo `value` e um valor.

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll('#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])');
assert(els1.length > 0 && els2.length === 0);
```

Você deve ter pelo menos um elemento `textarea` aninhado dentro de `#survey-form`.

```js
const el = document.querySelector('#survey-form textarea');
assert(!!el);
```

Você deve ter um elemento `input` ou `button` com um `id` de `submit`.

```js
const el = document.getElementById('submit');
assert(!!el && (el.tagName === 'INPUT' || el.tagName === 'BUTTON'));
```

O elemento `#submit` deve ter um `type` de `submit`.

```js
const el = document.getElementById('submit');
assert(!!el && el.type === 'submit');
```

O elemento `#submit` deve estar aninhado dentro de `#survey-form`.

```js
const el = document.querySelector('#survey-form #submit');
assert(!!el);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Survey Form</title>
</head>

<body>

</body>

</html> 
```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="container">
      <header class="header">
        <h1 id="title" class="text-center">freeCodeCamp Survey Form</h1>
        <p id="description" class="description text-center">
          Thank you for taking the time to help us improve the platform
        </p>
      </header>
      <form id="survey-form">
        <div class="form-group">
          <label id="name-label" for="name">Name</label>
          <input
            type="text"
            name="name"
            id="name"
            class="form-control"
            placeholder="Enter your name"
            required
          />
        </div>
        <div class="form-group">
          <label id="email-label" for="email">Email</label>
          <input
            type="email"
            name="email"
            id="email"
            class="form-control"
            placeholder="Enter your Email"
            required
          />
        </div>
        <div class="form-group">
          <label id="number-label" for="number"
            >Age<span class="clue">(optional)</span></label
          >
          <input
            type="number"
            name="age"
            id="number"
            min="10"
            max="99"
            class="form-control"
            placeholder="Age"
          />
        </div>
        <div class="form-group">
          <p>Which option best describes your current role?</p>
          <select id="dropdown" name="role" class="form-control" required>
            <option disabled selected value>Select current role</option>
            <option value="student">Student</option>
            <option value="job">Full Time Job</option>
            <option value="learner">Full Time Learner</option>
            <option value="preferNo">Prefer not to say</option>
            <option value="other">Other</option>
          </select>
        </div>

        <div class="form-group">
          <p>Would you recommend freeCodeCamp to a friend?</p>
          <label>
            <input
              name="user-recommend"
              value="definitely"
              type="radio"
              class="input-radio"
              checked
            />Definitely</label
          >
          <label>
            <input
              name="user-recommend"
              value="maybe"
              type="radio"
              class="input-radio"
            />Maybe</label
          >

          <label
            ><input
              name="user-recommend"
              value="not-sure"
              type="radio"
              class="input-radio"
            />Not sure</label
          >
        </div>

        <div class="form-group">
          <p>What is your favorite feature of freeCodeCamp?</p>
          <select id="most-like" name="mostLike" class="form-control" required>
            <option disabled selected value>Select an option</option>
            <option value="challenges">Challenges</option>
            <option value="projects">Projects</option>
            <option value="community">Community</option>
            <option value="openSource">Open Source</option>
          </select>
        </div>

        <div class="form-group">
          <p>
            What would you like to see improved?
            <span class="clue">(Check all that apply)</span>
          </p>

          <label
            ><input
              name="prefer"
              value="front-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Front-end Projects</label
          >
          <label>
            <input
              name="prefer"
              value="back-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Back-end Projects</label
          >
          <label
            ><input
              name="prefer"
              value="data-visualization"
              type="checkbox"
              class="input-checkbox"
            />Data Visualization</label
          >
          <label
            ><input
              name="prefer"
              value="challenges"
              type="checkbox"
              class="input-checkbox"
            />Challenges</label
          >
          <label
            ><input
              name="prefer"
              value="open-source-community"
              type="checkbox"
              class="input-checkbox"
            />Open Source Community</label
          >
          <label
            ><input
              name="prefer"
              value="gitter-help-rooms"
              type="checkbox"
              class="input-checkbox"
            />Gitter help rooms</label
          >
          <label
            ><input
              name="prefer"
              value="videos"
              type="checkbox"
              class="input-checkbox"
            />Videos</label
          >
          <label
            ><input
              name="prefer"
              value="city-meetups"
              type="checkbox"
              class="input-checkbox"
            />City Meetups</label
          >
          <label
            ><input
              name="prefer"
              value="wiki"
              type="checkbox"
              class="input-checkbox"
            />Wiki</label
          >
          <label
            ><input
              name="prefer"
              value="forum"
              type="checkbox"
              class="input-checkbox"
            />Forum</label
          >
          <label
            ><input
              name="prefer"
              value="additional-courses"
              type="checkbox"
              class="input-checkbox"
            />Additional Courses</label
          >
        </div>

        <div class="form-group">
          <p>Any comments or suggestions?</p>
          <textarea
            id="comments"
            class="input-textarea"
            name="comment"
            placeholder="Enter your comment here..."
          ></textarea>
        </div>

        <div class="form-group">
          <button type="submit" id="submit" class="submit-button">
            Submit
          </button>
        </div>
      </form>
    </div>
  </body>
</html>
```
