---
id: 67051431a73f1ca25d9a6f25
title: Construir um Quadro de Ideias para Projetos
challengeType: 26
dashedName: build-a-project-idea-board
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório. 

**Histórias de Usuário:**

1. Você deve definir uma constante de objeto chamada `projectStatus` com as três chaves: `PENDING`, `SUCCESS` e `FAILURE`. Cada status deve ser atribuído a um objeto com uma chave `description` com os valores `Execução Pendente`, `Executado com Sucesso` e `Falha na Execução`, respectivamente.
2. Você deve definir uma classe chamada `ProjectIdea` com um `constructor` que recebe `title` e `description` como parâmetros. Inicialize as propriedades `title` e `description` com os parâmetros fornecidos. A classe também deve ter uma propriedade chamada `status` que é definida com o valor `projectStatus.PENDING` por padrão.
3. Você deve definir um método chamado `updateProjectStatus` dentro da classe `ProjectIdea`. Este método deve aceitar um parâmetro `newStatus` e atualizar a propriedade `status` para o valor fornecido.
4. Você deve definir uma classe `ProjectIdeaBoard` com um `constructor` que aceita um `title` e inicializa um array vazio chamado `ideas` para armazenar instâncias da classe `ProjectIdea`.
5. Você deve definir um método chamado `pin` dentro da classe `ProjectIdeaBoard` que aceite uma instância da classe `ProjectIdea` e adicione a instância fornecida ao array `ideas`.
6. Você deve definir um método chamado `unpin` dentro da classe `ProjectIdeaBoard`. Este método deve aceitar uma instância da classe `ProjectIdea` e removê-la do array `ideas`.
7. Você deve definir um método chamado `count` que retorne o número de ideias de projeto no array `ProjectIdeaBoard` fornecido.
8. Você deve definir um método chamado `formatToString` que retorne o nome das ideias dos projetos, sua descrição e status no formato:

```js
<ProjectIdeaBoard title> has <ProjectIdeaBoard count> idea(s)
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
.
.
.
```

# --hints--

Você deve definir uma constante de objeto chamada `projectStatus` com os status e descrições corretos do projeto.

```js
assert.isObject(projectStatus);
assert.sameMembers(Object.keys(projectStatus), ["PENDING", "SUCCESS", "FAILURE"]);
assert.deepEqual(projectStatus.PENDING, { description: "Pending Execution" });
assert.deepEqual(projectStatus.SUCCESS, { description: "Executed Successfully" });
assert.deepEqual(projectStatus.FAILURE, { description: "Execution Failed" });
```

Você deve ter uma classe `ProjectIdea`.

```js
assert.isFunction(ProjectIdea);
assert.isObject(new ProjectIdea());
```

Sua classe `ProjectIdea` deve inicializar a propriedade `title` e a propriedade `description` com base nos parâmetros passados.Também deve definir o status para `projectStatus.PENDING`.

```js
const test = new ProjectIdea("sample title", "sample description");
assert.isAbove(test.title.length, 0);
assert.isAbove(test.description.length, 0);
assert.strictEqual(test.status.description, "Pending Execution");
```

Sua classe `ProjectIdea` deve ter um método `updateProjectStatus`.

```js
const projectIdea = new ProjectIdea();
assert.isFunction(projectIdea.updateProjectStatus);
```

Você deve ter uma classe `ProjectIdeaBoard`.

```js
assert.isFunction(ProjectIdeaBoard);
assert.isObject(new ProjectIdeaBoard());
```

Seu `ProjectIdeaBoard` deve inicializar a propriedade `title` com base no parâmetro passado e inicializar um array vazio chamado `ideas` para armazenar instâncias da classe `ProjectIdea`.

```js
const test = new ProjectIdeaBoard("Sample Title");
assert.isArray(test.ideas);
assert.lengthOf(test.ideas, 0);
assert.exists(test.title);
```

Sua classe `ProjectIdeaBoard` deve ter um método `pin`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.pin);
```

Sua classe `ProjectIdeaBoard` deve ter um método `unpin`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.unpin);
```

Sua classe `ProjectIdeaBoard` deve ter um método `count`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.count);
```

Sua classe `ProjectIdeaBoard` deve ter um método `formatToString`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.formatToString);
```

`new ProjectIdea("Trancas Inteligentes para Janelas", "Um projeto de automação que permite aos usuários trancar e destrancar janelas automaticamente com base nas condições climáticas.")` deve retornar `{ title: 'Trancas Inteligentes para Janelas', description: 'Um projeto de automação que permite aos usuários trancar e destrancar janelas automaticamente com base nas condições climáticas.', status: { description: 'Execução Pendente' } }`.

```js
const windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);
assert.deepEqual(windowLocks, {
  title: "Smart Window Locks",
  description:
    "An automation project allowing users to lock, unlock windows automatically based on weather conditions.",
  status: { description: "Pending Execution" },
});
```

Chamar `updateProjectStatus(projectStatus.SUCCESS)` em `new ProjectIdea("Fitness Tracker App", "Um aplicativo que monitora os treinos, dieta e padrões de sono do usuário.")` deve atualizar o status para `{ description: 'Executado com Sucesso' }`.

```js
let mobileApp = new ProjectIdea(
  "Fitness Tracker App",
  "An app that tracks user workouts, diet, and sleep patterns."
);

// Update the project status to SUCCESS
mobileApp.updateProjectStatus(projectStatus.SUCCESS);
assert.deepEqual(mobileApp, {
  title: "Fitness Tracker App",
  description: "An app that tracks user workouts, diet, and sleep patterns.",
  status: { description: "Executed Successfully" },
});
```

Chamar `updateProjectStatus(projectStatus.FAILURE)` em `new ProjectIdea("Robô Chef de Café da Manhã", "Um robô que pode seguir uma lista de instruções fornecida e preparar o café da manhã para o usuário, informando-o através do telefone.")` deve atualizar o status para `{ description: 'Execução Falhou' }`.

```js
let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);

// Update the project status to FAILURE
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

assert.deepEqual(breakfastChef, {
  title: "Breakfast Chef Robot",
  description:
    "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.",
  status: { description: "Execution Failed" },
});
```

Chamar `updateProjectStatus(projectStatus.SUCCESS)` em `new ProjectIdea("Loja Online de Jogos Usados", "Uma plataforma online onde os usuários podem comprar cópias físicas de jogos usados de outros usuários.")` deve atualizar o status para `{ description: 'Executado com Sucesso' }`.

```js
let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);

// Update the project status set to SUCCESS
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

assert.deepEqual(usedVideoGamesStore, {
  title: "Online Used Video Games Store",
  description:
    "An online platform where users can buy second hand physical copies of video games from other users.",
  status: { description: "Executed Successfully" },
});
```

Você deve ser capaz de fixar um objeto `ProjectIdea` no seu `ProjectIdeaBoard` usando o método `pin`.

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin the idea
let techProjects = new ProjectIdeaBoard("Tech Projects Board");
techProjects.pin(smartHome);

// Correct the assertion to compare the first idea in the array
assert.deepEqual(techProjects.ideas[0], {
    title: 'Smart Home System',
    description: 'An integrated system to control lighting, temperature, and security devices remotely.',
    status: { description: 'Pending Execution' }
});
```

Você deve ser capaz de desfixar um objeto `ProjectIdea` do seu `ProjectIdeaBoard` usando o método `unpin`.

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin ideas
let techProjects = new ProjectIdeaBoard("Tech Projects Board");

techProjects.pin(smartHome);
techProjects.unpin(smartHome);

assert.isEmpty(techProjects.ideas);
```

Quando `new ProjectIdeaBoard("Empty Board")` estiver vazio, `emptyBoard.formatToString()` deve retornar `Empty Board has 0 idea(s)\n`.

```js

let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin ideas
let techProjects = new ProjectIdeaBoard("Tech Projects Board");

// Create an empty board
let emptyBoard = new ProjectIdeaBoard("Empty Board");

// Attempt to unpin a project from an empty board
emptyBoard.unpin(smartHome);

assert.equal(emptyBoard.formatToString(), "Empty Board has 0 idea(s)\n");
```

Quando você fixa `new ProjectIdea("Sistema de Casa Inteligente", "Um sistema integrado para controlar remotamente iluminação, temperatura e dispositivos de segurança.")` em `new ProjectIdeaBoard("Quadro de Projetos Tecnológicos")`, `techProjects.formatToString()` deve retornar `Quadro de Projetos Tecnológicos tem 1 ideia(s)\nSistema de Casa Inteligente (Execução Pendente) - Um sistema integrado para controlar remotamente iluminação, temperatura e dispositivos de segurança.\n`.

```js
let techProjectsBoard = new ProjectIdeaBoard("Tech Projects Board");
let smartHome = new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.");
techProjectsBoard.pin(smartHome);

assert.equal(techProjectsBoard.formatToString(), "Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n"); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const projectStatus = {
  PENDING: { description: "Pending Execution" },
  SUCCESS: { description: "Executed Successfully" },
  FAILURE: { description: "Execution Failed" },
};

class ProjectIdea {
  constructor(title, description) {
    this.title = title;
    this.description = description;
    this.status = projectStatus.PENDING;
  }

  updateProjectStatus(newStatus) {
    this.status = newStatus;
  }
}

class ProjectIdeaBoard {
  constructor(title) {
    this.title = title;
    this.ideas = [];
  }

  pin(newIdea) {
    this.ideas.push(newIdea);
  }

  unpin(idea) {
    this.ideas.splice(this.ideas.indexOf(idea), 1);
  }

  count() {
    return this.ideas.length;
  }

  formatToString() {
    let representation = `${this.title} has ${this.count()} idea(s)\n`;

    this.ideas.forEach(
      (idea) =>
        (representation += `${idea.title} (${idea.status.description}) - ${idea.description}\n`),
    );
    return representation;
  }
}

let windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);

let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

let hardwareProjectIdeas = new ProjectIdeaBoard("Hardware Project Board");
hardwareProjectIdeas.pin(windowLocks);
hardwareProjectIdeas.pin(breakfastChef);

let softwareProjectIdeas = new ProjectIdeaBoard("Software Project Board");
softwareProjectIdeas.pin(usedVideoGamesStore);

console.log(softwareProjectIdeas.formatToString());

console.log(hardwareProjectIdeas.formatToString());
hardwareProjectIdeas.unpin(breakfastChef);
console.log(hardwareProjectIdeas.formatToString());
```
