---
id: 694a9ff65d4472dd451d8255
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

このワークショップでは、Pythonでブール値と条件文を使う練習を、映画チケット予約の計算機を作りながら行います。

まず、ユーザーと映画の上映に関する情報をストアする変数を宣言してください。

映画チケットの基本料金をストアする`base_price`という名前の変数を作成し、その値を`15`に設定してください。さらに、ユーザーの年齢をストアする`age`という名前の変数を作成し、その値を`21`に設定してください。

# --hints--

`base_price`という名前の変数を宣言してください。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("base_price")`)) })
```

整数値`15`を`base_price`に代入してください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("base_price").is_equivalent("base_price = 15")
`)) })
```

`age`という名前の変数を宣言してください。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("age")`)) })
```

整数値`21`を`age`に代入してください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("age").is_equivalent("age = 21")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
