---
id: 694cc1b515976889afd4cded
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

今度は映画のチケットに関するいくつかの文字列の値をストアする必要があります。Pythonではこのように文字列をストアできます。

```py
name = 'Alice'
```

ユーザーが選択した座席の種類をストアするために、`seat_type` という名前の変数を作成し、その値を文字列の `Gold` に設定してください。映画の上映時間をストアするために、`show_time` という名前の別の変数を作成し、その値を文字列の `Evening` に設定してください。

両方の値はシングルクォートか二重引用符で囲むことを忘れないでください。

# --hints--

`seat_type` という名前の変数を宣言してください。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("seat_type")`)) })
```

文字列の `Gold` を `seat_type` に代入してください。値はシングルクォートか二重引用符で囲むことを忘れないでください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("seat_type").is_equivalent("seat_type = 'Gold'")
`)) })
```

`show_time` という名前の変数を宣言してください。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("show_time")`)) })
```

文字列の `Evening` を `show_time` に代入してください。値はシングルクォートか二重引用符で囲むことを忘れないでください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("show_time").is_equivalent("show_time = 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
--fcc-editable-region--

--fcc-editable-region--
```
