---
id: 6958de0d773289a89515a6b8
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

今度は、ユーザーの年齢に基づいて夜のショーの予約が許可されているかどうかを確認します。

`if` 文を作成して、`age` が `21` 以上かどうかをチェックしてください。`if` 文の中で、`User is eligible for Evening shows` を出力してください。

# --hints--

コードには `if` 文が2つあるはずです。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1]`)) })
```

2つ目の `if` 文の条件は `age >= 21` にしてください。

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[1].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

新しい `User is eligible for Evening shows` 文の中で `if` を出力してください。メッセージはシングルクォートか二重引用符で囲むことを忘れないでください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[0].has_call("print('User is eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

--fcc-editable-region--

--fcc-editable-region--
```
