---
id: 6958dd85587e10c99162dd81
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

前のレッスンで学んだように、Pythonでは、`if` 文を使って条件が真の場合にコードを実行できます。

`if` 文は、`if` キーワードに続いて条件とコロンが続きます。条件が真の場合に実行されるコードはインデントされていなければならず、これを`if` 文の本体と呼びます。

```py
if condition:
    # Code to execute if condition is True
```

このステップでは、ユーザーの年齢に基づいて映画のチケットを予約できるかどうかを確認します。

`if` 文を作成して、`age` が `17` より大きいかどうかをチェックしてください。`if` 文の本体の中で、`User is eligible to book a ticket` を出力してください。これにより、ユーザーの年齢が `17` より大きい場合にのみメッセージが出力されます。

`if` 文の本体をインデントし、`print()` 呼び出しの中でメッセージをシングルクォートまたは二重引用符で囲むことを忘れないでください。

# --hints--

`if` 文が必要です。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[0]`)) })
```

`if` 文の条件は `age > 17` であるべきです。

```js 
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("age > 17") or _cond.is_equivalent("17 < age")
`)) })
```

`print()` 文の中に `if` 呼び出しが必要です。`if` 文の本体をインデントすることを忘れないでください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].find_calls("print")[0]
`)) })
```

`User is eligible to book a ticket` 文の中で `if` を出力してください。メッセージはシングルクォートまたは二重引用符で囲むことを忘れないでください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].has_call("print('User is eligible to book a ticket')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

--fcc-editable-region--

--fcc-editable-region--
```
