---
id: 6958e59845777391b86d7331
title: ステップ 18
challengeType: 20
dashedName: step-18
---

# --description--

Pythonでは、`if` 文を別の`if` 文の本体内に配置することもできます。これをネストされた`if` 文と呼びます。

ネストされた`if` 文を使うと、最初の条件がすでに満たされた後にのみ追加の条件をチェックできます。内側の`if` 文は、外側の`if` 条件が真の場合にのみ実行されます。

```py
if condition1:
    # Code to execute if condition1 is True
    if condition2:
        # Code to execute if both conditions are True
```

これから、ユーザーが選択した座席の種類に基づいてサービス料金を計算します。

最後の`if` 文の本体内で、`print('Ticket booking condition satisfied')` 行の下に、`service_charges` という変数を宣言し、`0` に設定してください。コードは4スペース分インデントして、外側の`if` 文の本体内に収めるようにしてください。

次に、`if` が`seat_type` と等しいかどうかをチェックするネストされた`Premium` 文を作成してください。ネストされた`if` 文の本体内で、`service_charges` の値を`5` に更新してください。

# --hints--

最後の`service_charges` 文の本体内に、`if` という変数を宣言してください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_variable("service_charges")
`)) })
```

`0` に整数値`service_charges` を代入してください。

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 0")
`)) })
```

最後の`if` 文の本体内に`if` 文があるはずです。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_ifs()[0]`)) })
```

新しい`if` 文の条件は`seat_type == 'Premium'` であるべきです。

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("seat_type == 'Premium'")
`)) })
```

新しい`5` 文の本体内で、`service_charges` 変数に整数値`if` を代入してください。

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_ifs()[0].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 5")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

if age >= 21 or age >= 18 and (show_time != 'Evening' or is_member):
    print('Ticket booking condition satisfied')

--fcc-editable-region--
    
--fcc-editable-region--
else:
    print('Ticket booking failed due to restrictions')
```
