---
id: 6958dfa0c74b237e1b56ce04
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

ユーザーがメンバーであれば、メンバーシップ割引の対象になります。

`discount` という名前の変数を作成し、その値を `0` に設定してください。これはユーザーが映画のチケットで受ける割引をストアします。

# --hints--

`discount` という名前の変数を宣言してください。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("discount")`)) })
```

整数値 `0` を `discount` に代入してください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("discount").is_equivalent("discount = 0")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

--fcc-editable-region--

--fcc-editable-region--
```
