---
id: 6958e1875e592b58e8852757
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

`extra_charges` という名前の変数を作成し、`0` を代入してください。これは週末に映画のチケットに追加される料金を表します。

次に、`if` が真であるかどうかを確認する `is_weekend` 文を作成してください。`if` 文の中で、`extra_charges` の値を `2` に更新し、ターミナルに `Extra charges will be applied` を出力してください。

# --hints--

`extra_charges` という名前の変数を宣言してください。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("extra_charges")`)) })
```

`0` に整数値 `extra_charges` を代入してください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("extra_charges").is_equivalent("extra_charges = 0")
`)) })
```

4つ目の `if` 文を作成してください。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3]`)) })
```

4つ目の `if` 文の条件は `is_weekend` にしてください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_conditions()[0].is_equivalent("is_weekend")
`)) })
```

新しい `2` 文の中で `extra_charges` 変数に整数値 `if` を代入してください。

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[3].find_bodies()[0].find_variable("extra_charges") 
assert _var.is_equivalent("extra_charges = 2")
`)) })
```

新しい `Extra charges will be applied` 文の中で `if` を出力してください。メッセージはシングルクォートか二重引用符で囲むことを忘れないでください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[0].has_call("print('Extra charges will be applied')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

--fcc-editable-region--

--fcc-editable-region--
```
