---
id: 6958e2377cacbc46a7f7383a
title: ステップ 13
challengeType: 20
dashedName: step-13
---

# --description--

今、`else` 文に `if` 節を追加し、`No extra charges will be applied` ブロック内で `else` を出力してください。これにより、追加料金の条件が false の場合にのみメッセージが出力されます。

次に、`else` 節の下に `print()` 呼び出しを使って、`Extra charges:` に続いて更新された `extra_charges` の値を表示するメッセージを出力し、ターミナルで結果を確認してください。

# --hints--

`else` 節があるはずです。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3].find_conditions()[1].is_empty()`)) })
```

`No extra charges will be applied` 節の中で `else` を出力してください。メッセージはシングルクォートか二重引用符で囲むことを忘れないでください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[1].has_call("print('No extra charges will be applied')")
`)) })
```

`print('Extra charges:', extra_charges)` 文の下に `if...else` があるはずです。

```js
({ test: () => (runPython(`
if_stmt = """if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')"""
print_call = "print('Extra charges:', extra_charges)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
--fcc-editable-region--

--fcc-editable-region--
```
