---
id: 6958e0ecc121b0ae3dd9c5e0
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

ユーザーがメンバーシップ割引の対象外の場合の処理も行う必要があります。

`else` 文に `if` 節を追加し、`User does not qualify for membership discount` ブロック内で `else` を出力してください。

変数に続けてメッセージを出力するには、次のようにします。

```py
print('Total:', total)
```

また、`discount` の更新された値も表示したいです。`if...else` 文の下に、`print()` 呼び出しを使って、`Discount:` と更新された `discount` の値を表示するメッセージを出力してください。その後、ターミナルの出力を確認してください。

# --hints--

`else` 節が必要です。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2].find_conditions()[1].is_empty()`)) })
```

`User does not qualify for membership discount` 節内で `else` を出力してください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[1].has_call("print('User does not qualify for membership discount')")
`)) })
```

`print('Discount:', discount)` 文の下に `if...else` が必要です。

```js
({ test: () => (runPython(`
if_stmt = """if is_member:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')"""
print_call = "print('Discount:', discount)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
if is_member:
    discount = 3
    print('User qualifies for membership discount')
--fcc-editable-region--

--fcc-editable-region--
```
