---
id: 696099c1ac72a4e585de1d5b
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

ある情報は真か偽のどちらかしかありません。前のレッスンで学んだように、これはブール値を使って表せます。

ユーザーがメンバーかどうかを示すために、`is_member` という名前の変数を作成し、その値を `True` に設定してください。

`is_member` 変数の下に、映画の上映が週末かどうかを示すために `is_weekend` という名前の別の変数を作成し、その値を `False` に設定してください。値は引用符で囲まないでください。

# --hints--

`is_member` という名前の変数を宣言してください。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_member")`)) })
```

ブール値 `True` を `is_member` に代入してください。値は引用符で囲まないでください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = True")
`)) })
```

`is_weekend` という名前の変数を宣言してください。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_weekend")`)) })
```

ブール値 `False` を `is_weekend` に代入してください。値は引用符で囲まないでください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_weekend").is_equivalent("is_weekend = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--

--fcc-editable-region--
```
