---
id: 6960a5165f7efe52a0c0aa39
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

Pythonでは、`and`演算子を使って複数の条件がすべて真かどうかを確認します。`if`文で2つの条件を組み合わせる方法は次の通りです。

```py
if condition1 and condition2:
    # Code to execute if all conditions are True
```

会員割引は、`age`が`21`以上の場合にのみ会員に適用されるべきです。

`if is_member:`行の条件を更新して、`and`演算子を使い、既存の条件と`age`が`21`以上かどうかを確認する別の条件を組み合わせてください。

# --hints--

3つ目の`if`文の条件は`is_member and age >= 21`にしてください。

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[2].find_conditions()[0] 
assert _cond.is_equivalent("is_member and age >= 21") or _cond.is_equivalent("is_member and 21 <= age")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--
if is_member:
--fcc-editable-region--
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
