---
id: 696099abb6acc539b27055a1
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

前のステップでは、`if` 文を使って条件をチェックしました。`else` 節は、その条件が false の場合を処理するために使われます。`if...else` 文の構文は次の通りです。

```py
if condition:
    # Code to execute if condition is True
else:
    # Code to execute if condition is False
```

次に、`else` 文に `if` 節を追加し、`User is not eligible for Evening shows` ブロック内で `else` を出力してください。これにより、`if` 文の条件が false のときだけ出力されます。

# --hints--

`else` 節があるはずです。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1].find_conditions()[1].is_empty()`)) })
```

`User is not eligible for Evening shows` 節の中で `else` を出力してください。メッセージはシングルクォートか二重引用符で囲むことを忘れないでください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[1].has_call("print('User is not eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
--fcc-editable-region--

--fcc-editable-region--
```
