---
id: 6960ba6e6c95ed4ad896c6f2
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

予約ルールにはもう一つ例外があります。`18`から`21`の間のユーザーは、メンバーであれば夜のショーを予約できます。

`if`文で複数の論理演算子が使われている場合、`and`で結合された条件は`or`で結合された条件より先に評価されます。Pythonでは丸括弧`()`を使って条件をグループ化し、評価の順序を制御します。

```py
if condition1 and (condition2 or condition3):
    # Code to execute if conditions evaluate to True
else:
    # Code to execute if conditions evaluate to False
```

`if age >= 21 or age >= 18 and show_time != 'Evening':`行の条件を更新して、`or`演算子を使い`is_member`が真であるかどうかをチェックする別の条件を追加してください。上記の例のように、`()`と`show_time != 'Evening'`の条件を丸括弧`is_member`でグループ化してください。

# --hints--

最後の`if`文の条件は`age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)`にしてください。

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--
if age >= 21 or age >= 18 and show_time != 'Evening':
--fcc-editable-region--
    print('Ticket booking condition satisfied')
else:
    print('Ticket booking failed due to restrictions')
```
