---
id: 6960b308bf41cc39c9b4b438
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

Pythonでは、`or`演算子を使って、2つの条件のうち少なくとも1つが真かどうかを確認します。`if`文での使い方は次の通りです。

```py
if condition1 or condition2:
    # Code to execute if any condition is True
```

ショーが夜の場合にも追加料金が適用されるべきです。`if is_weekend:`行の条件を更新して、`or`演算子を使い、既存の条件と`show_time`が文字列`Evening`と等しいかどうかをチェックする2つ目の条件を組み合わせてください。

# --hints--

最後の`if`文の条件は`is_weekend or show_time == 'Evening'`にしてください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_conditions()[0].is_equivalent("is_weekend or show_time == 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
--fcc-editable-region--
if is_weekend:
--fcc-editable-region--
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)
```
