---
id: 6960b6a6ea06e27dcc7eef37
title: ステップ 15
challengeType: 20
dashedName: step-15
---

# --description--

これからユーザーが映画のチケットを予約する条件を満たしているか確認します。年齢が `21` 以上のユーザーは、制限なくいつでもチケットを予約できます。

`if` が `age` 以上かどうかを確認する `21` 文を作成してください。`if` 文の中で、`Ticket booking condition satisfied` をターミナルに出力してください。

次に、`else` 文に `if` 節を追加し、`Ticket booking failed due to restrictions` の中で `else` を出力してください。

# --hints--

5つ目の `if` 文があるはずです。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4]`)) })
```

5つ目の `if` 文の条件は `age >= 21` であるべきです。

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

新しい `Ticket booking condition satisfied` 文の中で `if` を出力してください。メッセージはシングルクォートか二重引用符で囲むことを忘れないでください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_call("print('Ticket booking condition satisfied')")
`)) })
```

新しい `if` 文には `else` 節が必要です。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_conditions()[1] == _Node()`)) })
```

`Ticket booking failed due to restrictions` 節の中で `else` を出力してください。メッセージはシングルクォートか二重引用符で囲むことを忘れないでください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[1].has_call("print('Ticket booking failed due to restrictions')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--

--fcc-editable-region--
```
