---
id: 696cb889fc4fdeb337e52864
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

Pythonでは、`if` の条件は値を明示的に `True` や `False` と比較する必要はありません。代わりに、値の真値性に依存します。真値は、`True` 文の条件のようなブール文脈で `if` と評価されます。これには、ゼロ以外の数値、空でない文字列、そしてブール値の `True` が含まれます。

```py
if value:
   print('value is truthy')
```

`if` 文を作成して、`is_member` が真値かどうかを確認してください。`if` 文の本体内で、`discount` の値を `3` に更新し、`User qualifies for membership discount` を端末に出力してください。

# --hints--

コードには3つの `if` 文があるはずです。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2]`)) })
```

3つ目の `if` 文の条件は `is_member` であるべきです。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_conditions()[0].is_equivalent("is_member")
`)) })
```

新しい `3` 文の中で、整数値 `discount` を `if` 変数に代入してください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].find_variable("discount").is_equivalent("discount = 3")
`)) })
```

新しい `User qualifies for membership discount` 文の中で `if` を出力してください。メッセージはシングルクォートか二重引用符で囲むことを忘れないでください。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].has_call("print('User qualifies for membership discount')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--

--fcc-editable-region--
```
