---
id: 5dc174fcf86c76b9248c6eb2
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML-Elemente haben öffnende Tags wie `<h1>` und schließende Tags wie `</h1>`. Der Text, den ein Element anzeigt, steht zwischen seinem öffnenden und schließenden Tag.

Ändern Sie den Text des `h1`-Elements unten von `Hello World` zu `CatPhotoApp` und beobachten Sie die Änderung in der Browser-Vorschau.

Wenn Sie fertig sind, drücken Sie die Schaltfläche „Check Your Code“, um zu sehen, ob es korrekt ist.

# --hints--

Der Text `CatPhotoApp` sollte im Code vorhanden sein. Überprüfen Sie gegebenenfalls Ihre Rechtschreibung.

```js
assert(code.match(/catphotoapp/i));
```

Ihr `h1`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Ihre `h1`-Tags sollten in Kleinbuchstaben geschrieben sein. Nach Konvention werden alle HTML-Tags in Kleinbuchstaben geschrieben.

```js
assert.notMatch(code, /<\/?H1>/);
```

Ihr `h1`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
assert(code.match(/<\/h1\>/));
```

Sie haben mehr als ein `h1`-Element. Entfernen Sie das zusätzliche `h1`-Element.

```js
assert(document.querySelectorAll('h1').length === 1);
```

Der Text Ihres `h1`-Elements sollte `CatPhotoApp` sein. Sie haben entweder den Text ausgelassen, einen Tippfehler oder er steht nicht zwischen dem öffnenden und schließenden Tag des `h1`-Elements.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

