---
id: 5dc1798ff86c76b9248c6eb3
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Die `h1` bis `h6` Überschriftenelemente werden verwendet, um die Wichtigkeit des darunterliegenden Inhalts zu kennzeichnen. Je niedriger die Zahl, desto höher die Wichtigkeit, daher haben `h2` Elemente weniger Bedeutung als `h1` Elemente.

```html
<h1>most important heading element</h1>
<h2>second most important heading element</h2>
<h3>third most important heading element</h3>
<h4>fourth most important heading element</h4>
<h5>fifth most important heading element</h5>
<h6>least important heading element</h6>
```

Verwenden Sie pro Seite nur ein `h1` Element und platzieren Sie Überschriften mit geringerer Wichtigkeit unter Überschriften mit höherer Wichtigkeit.

Fügen Sie unter dem `h1` Element ein `h2` Element mit diesem Text hinzu:

`Cat Photos`

# --hints--

Ihr `h1` Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Ihr `h1` Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<` Zeichen.

```js
assert(code.match(/<\/h1\>/));
```

Sie sollten nur ein `h1` Element haben. Entfernen Sie das zusätzliche.

```js
assert(
  document.querySelector('h1') && document.querySelectorAll('h1').length === 1
);
```

Der Text Ihres `h1` Elements sollte `CatPhotoApp` sein. Sie haben entweder den Text ausgelassen oder einen Tippfehler.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Ihre `h2` Tags sollten in Kleinbuchstaben sein. Nach Konvention werden alle HTML-Tags in Kleinbuchstaben geschrieben.

```js
assert.notMatch(code, /<\/?H2>/);
```

Ihr `h2` Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert(document.querySelector('h2'));
```

Ihr `h2` Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<` Zeichen.

```js
assert(code.match(/<\/h2\>/));
```

Der Text Ihres `h2` Elements sollte `Cat Photos` sein. Platzieren Sie nur den Text `Cat Photos` zwischen den öffnenden und schließenden `h2` Tags.

```js
assert(document.querySelector('h2').innerText.toLowerCase() === 'cat photos');
```

Ihr `h2` Element sollte unter dem `h1` Element sein. Das `h1` Element hat eine höhere Wichtigkeit und muss über dem `h2` Element stehen.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H1') < collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>
    
--fcc-editable-region--
  </body>
</html>
```

