---
id: 5dc174fcf86c76b9248c6eb2
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML Elemente haben öffnende Tags wie `<h1>` und schließende Tags wie `</h1>`. Der Text des Elementes, der angezeigt werden soll, gehört zwischen seinen öffnenden und schließenden Tags.

Ändere den Text des unten stehenden `h1`-Elements von `Hello World` zu `CatPhotoApp` und schau dir die Veränderung im Browserfenster an.

When you are done, press the "Check Your Code" button to see if it's correct.

# --hints--

The text `CatPhotoApp` should be present in the code. You may want to check your spelling.

```js
assert(code.match(/catphotoapp/i));
```

Dein `h1`-Element sollte einen öffnenden Tag haben. Opening tags have this syntax: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Your `h1` tags should be in lowercase. By convention, all HTML tags are written in lowercase.

```js
assert.notMatch(code, /<\/?H1>/);
```

Dein `h1`-Element sollte ein schließenden Tag haben. Schließende Tags haben ein `/` direkt nach dem `<` Zeichen.

```js
assert(code.match(/<\/h1\>/));
```

You have more than one `h1` element. Remove the extra `h1` element.

```js
assert(document.querySelectorAll('h1').length === 1);
```

Der Text deines `h1`-Elements sollte `CatPhotoApp` lauten. You have either omitted the text, have a typo, or it is not between the `h1` element's opening and closing tags.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

