---
id: 5dc17d3bf86c76b9248c6eb4
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Das `p`-Element wird verwendet, um auf Websites einen Absatz mit Text zu erstellen. Erstellen Sie unter Ihrem `p`-Element ein `h2`-Element und geben Sie ihm folgenden Text:

`Everyone loves cute cats online!`

# --hints--

Ihr `p`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben folgende Syntax: `<elementName>`.

```js
assert(document.querySelector('p'));
```

Ihre `p`-Tags sollten in Kleinbuchstaben geschrieben sein. Nach Konvention werden alle HTML-Tags in Kleinbuchstaben geschrieben.

```js
assert.notMatch(code, /<\/?P>/);
```

Ihr `p`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
assert(code.match(/<\/p\>/));
```

Der Text Ihres `p`-Elements sollte `Everyone loves cute cats online!` sein. Sie haben entweder den Text weggelassen oder einen Tippfehler.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

Ihr `p`-Element sollte unter dem `h2`-Element stehen. Sie haben sie in der falschen Reihenfolge.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>
    
--fcc-editable-region--
  </body>
</html>
```

