---
id: 5dc17dc8f86c76b9248c6eb5
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Kommentare ermöglichen es Ihnen, Nachrichten zu hinterlassen, ohne die Anzeige im Browser zu beeinflussen. Sie erlauben Ihnen auch, Code zu deaktivieren. Ein Kommentar in HTML beginnt mit `<!--`, enthält beliebig viele Textzeilen und endet mit `-->`. 

Hier ist ein Beispiel für einen Kommentar mit dem `TODO: Remove h1`:

```html
<!-- TODO: Remove h1 -->
```

Fügen Sie einen Kommentar über dem `p`-Element mit diesem Text ein:

`TODO: Add link to cat photos`

# --hints--

Ihr Kommentar sollte mit `<!--` beginnen. Ihnen fehlt ein oder mehrere Zeichen, die den Anfang eines Kommentars definieren.

```js
assert(code.match(/<!--/));
```

Ihr Kommentar sollte mit `-->` enden. Ihnen fehlt ein oder mehrere Zeichen, die das Ende eines Kommentars definieren.

```js
assert(code.match(/-->/));
```

Ihr Code sollte keine zusätzlichen öffnenden oder schließenden Kommentarzeichen enthalten. Sie haben ein zusätzliches `<!--` oder `-->`, das im Browser angezeigt wird.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/<!--/g).length < 2 && noSpaces.match(/-->/g).length < 2);
```

Ihr Kommentar sollte über dem `p`-Element stehen. Sie haben sie in der falschen Reihenfolge.

```js
assert.match(
  code.replace(/\s/g, ''),
  /<!--(.*?)--><p>everyonelovescutecatsonline!<\/p>/i
);
```

Ihr Kommentar sollte den Text `TODO: Add link to cat photos` enthalten.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
    <p>Everyone loves cute cats online!</p>
    
--fcc-editable-region--
  </body>
</html>
```

