---
id: 5dc2385ff86c76b9248c6eb7
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 hat einige Elemente, die verschiedene Inhaltsbereiche identifizieren. Diese Elemente machen Ihr HTML leichter lesbar und helfen bei der Suchmaschinenoptimierung (SEO) sowie der Barrierefreiheit.

Das `main`-Element wird verwendet, um den Hauptinhalt des Body-Bereichs eines HTML-Dokuments darzustellen. Der Inhalt innerhalb des `main`-Elements sollte einzigartig für das Dokument sein und nicht in anderen Teilen des Dokuments wiederholt werden.

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Identifizieren Sie den Hauptabschnitt dieser Seite, indem Sie ein `<main>`-öffnendes Tag vor dem `h1`-Element und ein `</main>`-schließendes Tag nach dem `p`-Element hinzufügen.

# --hints--

Ihr `main`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert(document.querySelector('main'));
```

Ihre `main`-Tags sollten in Kleinbuchstaben geschrieben sein. Nach Konvention werden alle HTML-Tags in Kleinbuchstaben geschrieben.

```js
const mainTagMatches = code.matchAll(/<\/?main>/gi);
for (const match of mainTagMatches) {
  const tag = match[0];
  assert.strictEqual(
    tag,
    tag.toLowerCase()
  );
}
```

Ihr `main`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
assert(code.match(/<\/main\>/));
```

Das öffnende Tag Ihres `main`-Elements sollte unter dem öffnenden Tag des `body`-Elements stehen. Sie haben sie in der falschen Reihenfolge.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

Das öffnende Tag Ihres `main`-Elements sollte über dem `h1`-Element stehen. Sie haben sie in der falschen Reihenfolge.

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('MAIN') < collection.indexOf('H1'));
```

Das schließende Tag Ihres `main`-Elements sollte unter dem `p`-Element stehen. Sie haben sie in der falschen Reihenfolge.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode.contains(pNode));
assert.match(pNode.textContent.toLowerCase(), /everyone loves cute cats online!/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
  
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    
--fcc-editable-region--
  </body>
</html>
```

