---
id: 5dc23991f86c76b9248c6eb8
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Im vorherigen Schritt hast du die `h1`, `h2`, Kommentar und `p`-Elemente innerhalb deines `main`-Elements platziert. Das nennt man *nesting* (englisch für "verschachteln"). Verschachtelte Elemente sollten zwei Leerzeichen weiter rechts als das Element, in dem sie verschachtelt sind, sein. Diesen Abstand nennt man Einrückung (englisch: indentation) und wird verwendet, um HTML lesbarer zu machen.

Hier ist ein Beispiel zu Verschachtelung und Einrückung:

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Das `h1`-Element, das `h2`-Element und der Kommentar sind im folgenden Code zwei Leerzeichen mehr eingerückt als das `main`-Element. Verwende die Leertaste auf deiner Tastatur, um zwei Leerzeichen vor dem `p`-Element hinzuzufügen, sodass es richtig eingerückt wird.

# --hints--

Dein Code sollte ein `h2`-Element mit dem Text `Cat Photos` enthalten. Du hast es möglicherweise aus Versehen gelöscht. Oder es fehlen sowohl der öffnende und der schließende Tag. Oder du hast den Text geändert.

```js
assert(
  document.querySelector('h2') &&
    code.match(/<\/h2\>/) &&
    document.querySelector('h2').innerText.toLowerCase() === 'cat photos'
);
```

Du sollst nicht die `ul` oder `li`-Elemente aus dem Beispiel hinzufügen.

```js
assert(
  !document.querySelector('ul') && !document.querySelector('li')
);
```

Du solltest die Einrückung in der Zeile mit deinem `h2`-Element nicht ändern. Das dazugehörige öffnende Tag sollte 6 Leerzeichen vom Anfang der Zeile beginnen. Du kannst diese Challenge neu starten, um die ursprüngliche Einrückung wiederherzustellen.

```js
assert(code.toLowerCase().match(/<\/h1\>\s*\n\s{6}<h2>/));
```

Dein Code sollte einen Kommentar haben. Du hast den Kommentar aus den letzten Schritten entfernt.

```js
assert(code.match(/<!--.*-->/));
```

Der Text deines Kommentars sollte `TODO: Add link to cat photos` sein. Ändere nicht den Text oder die Einrückung des Kommentars.

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i));
```

Du solltest die Einrückung in der Zeile mit deinem Kommentar-Element nicht ändern. Der öffnende Tag sollte 6 Leerzeichen vom Anfang der Zeile beginnen. Du kannst diese Challenge neu starten, um die ursprüngliche Einrückung wiederherzustellen.

```js
assert(
  code
    .toLowerCase()
    .match(/<\/h2>\s*\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/)
);
```

Dein Code sollte ein `p`-Element beinhalten. Du hast das `p`-Element aus einem vorherigen Schritt entfernt.

```js
assert(document.querySelector('p'));
```

Der Text deines `p`-Elements sollte `Everyone loves cute cats online!` sein. Ändere den Text, die Einrückung oder die Satzzeichen des `p`-Elements nicht.

```js
assert.match(
  document
    .querySelector('p')
    .innerText.toLowerCase(),
    /everyone\s+loves\s+cute\s+cats\s+online!$/
);
```

Der öffnende `p` Tag sollte die gleiche Einrückung, wie das `h2`-Element und der Kommentar haben. Der öffnende Tag sollte 6 Leerzeichen vom Anfang der Zeile beginnen.

```js
assert(code.toLowerCase().match(/-->\s*\n\s{6}<p>/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```

