---
id: 5dc24073f86c76b9248c6ebb
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML-<dfn>Attribute</dfn> sind spezielle Wörter, die im öffnenden Tag eines Elements verwendet werden, um das Verhalten des Elements zu steuern. Das `src`-Attribut in einem `img`-Element gibt die URL des Bildes an (wo sich das Bild befindet).

Hier ist ein Beispiel für ein `img`-Element mit einem `src`-Attribut, das auf das freeCodeCamp-Logo zeigt:

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

Fügen Sie im vorhandenen `img`-Element ein `src`-Attribut mit dieser URL hinzu:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

Ihr Code sollte ein `img`-Element enthalten. Möglicherweise haben Sie das `img`-Element entfernt oder den Wert des `src`-Attributs nicht in Anführungszeichen gesetzt.

```js
assert(document.querySelector('img'));
```

Ihr `img`-Element sollte ein `src`-Attribut haben. Sie haben entweder das Attribut weggelassen oder einen Tippfehler. Achten Sie darauf, dass zwischen dem Elementnamen und dem Attributnamen ein Leerzeichen steht.

```js
assert(document.querySelector('img').src);
```

Das `src`-Attribut Ihres `img`-Elements sollte auf '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`' gesetzt sein. Sie haben entweder die URL weggelassen oder einen Tippfehler. Die Groß- und Kleinschreibung der URL ist wichtig.

```js
assert.strictEqual(document.querySelector('img').src, 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

Obwohl Sie das `img`-Attribut des `src`-Elements auf die richtige URL gesetzt haben, wird empfohlen, den Wert eines Attributs immer in Anführungszeichen zu setzen.

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

Ihr `img`-Element sollte entweder mit `>` oder `/>` geschlossen werden.

```js
assert.match(code, /<img\s+src\s*=\s*("|')\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\s*\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```

