---
id: 5ddbd81294d8ddc1510a8e56
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Der Text eines Links muss zwischen dem öffnenden und schließenden Tag eines Ankerelements (`a`) stehen.

Hier ist ein Beispiel für einen Link mit dem Text `click here to go to freeCodeCamp.org`:

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

Fügen Sie den Ankertext `cat photos` zum Ankerelement hinzu. Dieser wird zum Text des Links.

# --hints--

Ihr Ankerelement (`a`) sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert(document.querySelector('a'));
```

Ihr Ankerelement (`a`) sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
assert(code.match(/<\/a\>/));
```

Der Text Ihres Ankerelements (`a`) sollte `cat photos` sein. Stellen Sie sicher, dass Sie den Linktext zwischen dem öffnenden und schließenden Tag des Ankerelements (`a`) platzieren.

```js
assert.strictEqual(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

