---
id: 5dfb6250eacea3f48c6300b2
title: Schritt 24
challengeType: 0
dashedName: step-24
---

# --description--

Fügen Sie nach der ungeordneten Liste ein neues Bild mit einem `src`-Attributwert hinzu:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

Und seinen `alt`-Attributwert auf:

`A slice of lasagna on a plate.`

# --hints--

Es sollte ein `img`-Element direkt nach dem schließenden `</ul>`-Tag stehen.

```js
assert.equal(document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.nodeName, 'IMG');
```

Das neue Bild hat kein `alt`-Attribut. Prüfen Sie, ob nach dem Namen des öffnenden Tags ein Leerzeichen steht und/oder vor allen Attributnamen Leerzeichen sind.

```js
assert.isTrue(document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.hasAttribute('alt'));
```

Das neue Bild sollte einen `alt`-Wert von `A slice of lasagna on a plate.` haben. Stellen Sie sicher, dass der Wert des `alt`-Attributs in Anführungszeichen steht.

```js
assert.match(
  document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.getAttribute('alt')?.replace(/\s+/g, ' '), /^A slice of lasagna on a plate\.?$/i 
);
```

Das neue Bild hat kein `src`-Attribut. Prüfen Sie, ob nach dem Namen des öffnenden Tags ein Leerzeichen steht und/oder vor allen Attributnamen Leerzeichen sind.

```js
assert.isTrue(document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.hasAttribute('src'));
```

Das neue Bild sollte einen `src`-Wert von `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg` haben. Stellen Sie sicher, dass der Wert des `src`-Attributs in Anführungszeichen steht.

```js
assert.strictEqual(
  document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.getAttribute('src'),
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

Obwohl Sie das `src`-Attribut des neuen Bilds auf die korrekte URL gesetzt haben, wird empfohlen, den Wert eines Attributs immer in Anführungszeichen zu setzen.

```js
assert.isTrue(!/\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

