---
id: 5ef9b03c81a63668521804dc
title: Schritt 46
challengeType: 0
dashedName: step-46
---

# --description--

Sie können Optionsschaltflächen für Fragen verwenden, bei denen nur eine Antwort aus mehreren Optionen ausgewählt werden soll.

Hier ist ein Beispiel für eine Optionsschaltfläche mit dem Text gesetzt als `cat`:

```html
<input type="radio"> cat
```

Denken Sie daran, dass ein `input`-Element ein Void-Element ist.

Fügen Sie vor der Texteingabe eine Optionsschaltfläche mit dem Text gesetzt als hinzu:

`Indoor`

# --hints--

Sie sollten ein `input`-Element für Ihre Optionsschaltfläche erstellen. Überprüfen Sie die Syntax.

```js
assert(document.querySelectorAll('form > input').length >= 2);
```

Ihr `input`-Element sollte ein öffnendes Tag haben, aber kein schließendes Tag.

```js
assert(document.querySelector('form > input') && !code.match(/<\/input\>/g));
```

Sie sollten nur ein `input`-Element für Ihre Optionsschaltfläche hinzugefügt haben. Entfernen Sie alle zusätzlichen.

```js
assert(document.querySelectorAll('form > input'), 2);
```

Ihr neues `input`-Element sollte über dem vorhandenen `input` mit dem `type`-Attribut gesetzt auf `text` stehen. Sie haben sie in der falschen Reihenfolge.

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

Ihr neues `input`-Element hat kein `type`-Attribut. Überprüfen Sie, ob nach dem Namen des öffnenden Tags ein Leerzeichen steht.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

Ihr neues `input`-Element sollte nur ein `type`-Attribut haben. Entfernen Sie alle zusätzlichen.

```js
assert(document.querySelector('input')
  ?.getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

Ihr neues `input`-Element sollte ein `type`-Attribut mit dem Wert `radio` haben. Sie haben entweder den Wert weggelassen oder einen Tippfehler. Denken Sie daran, dass Attributwerte in Anführungszeichen stehen sollten.

```js
assert.match(
  document.querySelector('input')?.getAttribute('type'),
  /^radio$/i
);
```

Obwohl Sie das `type`-Attribut des neuen `input`-Elements auf `radio` gesetzt haben, wird empfohlen, den Wert eines Attributs immer in Anführungszeichen zu setzen.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

Der `radio`-Button-Text ` Indoor` sollte sich nach dem Button befinden, nicht davor.

```js
const radioInputElem = document.querySelector('input');
assert.notMatch(radioInputElem?.previousSibling?.nodeValue, /Indoor/i);
```

Der Text ` Indoor` sollte sich direkt rechts neben Ihrem `radio`-Button befinden. Sie haben entweder den Text weggelassen oder einen Tippfehler.

```js
const radioInputElem = document.querySelector('input');
assert.match(radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '), /\s*Indoor/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
 
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

