---
id: 5ef9b03c81a63668521804d8
title: Schritt 39
challengeType: 0
dashedName: step-39
---

# --description--

Das `input`-Element bietet Ihnen mehrere Möglichkeiten, Daten aus einem Webformular zu sammeln. Wie `img`-Elemente sind `input`-Elemente ein <dfn>void element</dfn> und benötigen keine schließenden Tags.

Verschachteln Sie ein `input`-Element im `form`-Element.

# --hints--

Ihr `form`-Element sollte ein öffnendes Tag und ein schließendes Tag in der richtigen Reihenfolge haben. Möglicherweise fehlt Ihnen eines oder beide der erforderlichen Tags, oder sie sind in der falschen Reihenfolge.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

Das öffnende Tag Ihres `form`-Elements sollte nur ein `action`-Attribut enthalten. Entfernen Sie alles andere, was Sie möglicherweise eingegeben haben.

```js
assert([...document.querySelector('form').attributes].length < 2);
```

Sie sollten ein `input`-Element erstellen. Überprüfen Sie die Syntax.

```js
assert(document.querySelector('input'));
```

Ihr `input`-Element sollte ein öffnendes Tag haben, aber kein schließendes Tag.

```js
assert(document.querySelector('input') && !code.match(/<\/input\>/g));
```

Ihr `input`-Element sollte innerhalb des `form`-Elements verschachtelt sein.

```js
assert(document.querySelector('form > input'));
```

Ihr `form` sollte nur das `input`-Element enthalten. Entfernen Sie alle HTML-Elemente oder Texte zwischen den Tags des `form`-Elements.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

