---
id: 5ef9b03c81a63668521804e5
title: Schritt 63
challengeType: 0
dashedName: step-63
---

# --description--

Um ein Kontrollkästchen standardmäßig aktiviert oder eine Optionsschaltfläche standardmäßig ausgewählt zu machen, müssen Sie das `checked`-Attribut hinzufügen.

Hier ist ein Beispiel für eine Optionsschaltfläche mit dem `checked`-Attribut:

```html
<input checked type="radio" name="meal" value="breakfast"> Breakfast
```

Es ist nicht nötig, dem `checked`-Attribut einen Wert zuzuweisen. Fügen Sie stattdessen einfach das Wort `checked` zum `input`-Element hinzu, und achten Sie darauf, dass ein Leerzeichen zwischen diesem und anderen Attributen steht.

Machen Sie die erste Optionsschaltfläche und das erste Kontrollkästchen standardmäßig ausgewählt.

# --hints--

Stellen Sie sicher, dass weiterhin zwei Optionsschaltflächen und drei Kontrollkästchen in ihren jeweiligen `fieldset`-Elementen verschachtelt sind.

```js
assert.isTrue(
  document.querySelectorAll('input[type="radio"]').length === 2 &&
     document.querySelectorAll('fieldset > input[type="checkbox"]').length === 3
);
```

Ihre erste Optionsschaltfläche, bei der `id` auf `indoor` gesetzt ist, sollte das `checked`-Attribut haben.

```js
assert.isTrue(document.querySelector('input[type="radio"]')?.hasAttribute('checked'));
```

Die zweite Optionsschaltfläche sollte das `checked`-Attribut nicht haben.

```js
assert.isFalse(document.querySelectorAll('input[type="radio"]')?.[1]?.hasAttribute('checked'));
```

Dem ersten Kontrollkästchen fehlt das `checked`-Attribut.

```js
assert.isTrue(document.querySelector('input[type="checkbox"]')?.hasAttribute('checked'));
```

Das zweite Kontrollkästchen sollte das `checked`-Attribut nicht haben.

```js
assert.isFalse(document.querySelectorAll('input[type="checkbox"]')?.[1]?.hasAttribute('checked'));
```

Das dritte Kontrollkästchen sollte das `checked`-Attribut nicht haben.

```js
assert.isFalse(document.querySelectorAll('input[type="checkbox"]')?.[2]?.hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
            <input id="loving" type="checkbox" name="personality" value="loving"> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality" value="lazy"> <label for="lazy">Lazy</label>
            <input id="energetic" type="checkbox" name="personality" value="energetic"> <label for="energetic"> Energetic</label>
          </fieldset>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

