---
id: 5ef9b03c81a63668521804d9
title: Schritt 42
challengeType: 0
dashedName: step-42
---

# --description--

Platzhaltertext wird verwendet, um Menschen einen Hinweis darauf zu geben, welche Art von Informationen sie in eine Eingabe eingeben sollen.

Hier ist ein Beispiel für ein `input`-Element mit einem Platzhalter, der auf `e.g. Jane Doe` gesetzt ist:

```html
<input type="text" placeholder="e.g. Jane Doe">
```

Fügen Sie den Platzhaltertext `cat photo URL` zu Ihrem `input`-Element hinzu.

# --hints--

Sie haben entweder Ihr `input`-Element gelöscht oder es hat eine ungültige Syntax. Alle Werte von Attributen sollten in Anführungszeichen stehen.

```js
assert(document.querySelectorAll('input').length);
```

Ihr `form` sollte nur das `input`-Element enthalten. Entfernen Sie alle zusätzlichen HTML-Elemente oder Texte innerhalb des `form`-Elements.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Ihr `input`-Element hat kein `placeholder`-Attribut. Überprüfen Sie, ob nach dem Namen des öffnenden Tags ein Leerzeichen steht und/oder ob vor allen Attributnamen Leerzeichen sind.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('placeholder'));
```

Ihr `input`-Element sollte ein `placeholder`-Attribut mit dem Wert `cat photo URL` haben. Sie haben entweder den Wert weggelassen oder einen Tippfehler. Denken Sie daran, dass Attributwerte in Anführungszeichen stehen sollten.

```js
assert(
  document.querySelector('input')
    ?.getAttribute('placeholder')
    .replace(/\s+/g, ' ')
    .match(/^cat photo URL$/i)
);
```

Obwohl Sie das `placeholder`-Attribut des `input`-Elements auf `cat photo URL` gesetzt haben, wird empfohlen, den Wert eines Attributs immer in Anführungszeichen zu setzen.

```js
assert(!/\<\s*input\s+placeholder\s*=\s*cat\s+photo\s+url/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

