---
id: 5ef9b03c81a63668521804de
title: Schritt 50
challengeType: 0
dashedName: step-50
---

# --description--

Beachten Sie, dass beide Optionsschaltflächen gleichzeitig ausgewählt sein können. Damit das Auswählen einer Optionsschaltfläche automatisch die andere deselektiert, müssen beide Schaltflächen ein `name`-Attribut mit demselben Wert haben.

Hier ist ein Beispiel für zwei Optionsschaltflächen mit demselben `name`-Attribut:

```html
<input type="radio" name="meal"> Breakfast
<input type="radio" name="meal"> Lunch
```

Fügen Sie beiden Optionsschaltflächen das `name`-Attribut mit dem Wert `indoor-outdoor` hinzu.

# --hints--

Beide Optionsschaltflächen sollten sich weiterhin zwischen öffnenden und schließenden `label`-Element-Tags befinden.

```js
const labelChildNodes = [...document.querySelectorAll('form > label')].map(
  (node) => node.childNodes
);
assert(
  labelChildNodes.filter((childNode) => childNode[0].nodeName === 'INPUT')
    .length === 2
);
```

Beide Optionsschaltflächen sollten ein `name`-Attribut haben. Prüfen Sie, ob nach dem Namen des öffnenden Tags ein Leerzeichen ist und/oder vor allen Attributnamen Leerzeichen stehen.

```js
const radioButtons = [...document.querySelectorAll('input[type="radio"]')];
assert(radioButtons.every((btn) => btn.hasAttribute('name')));
```

Beide Optionsschaltflächen sollten ein `type`-Attribut haben. Prüfen Sie, ob nach dem Namen des öffnenden Tags ein Leerzeichen ist und/oder vor allen Attributnamen Leerzeichen stehen.

```js
const radioButtons = [...document.querySelectorAll('input')];
assert(radioButtons.every((btn) => btn.hasAttribute('type')));
```

Beide Optionsschaltflächen sollten ein `name`-Attribut mit dem Wert `indoor-outdoor` haben. Sie haben entweder den Wert weggelassen oder einen Tippfehler. Denken Sie daran, dass Attributwerte in Anführungszeichen stehen sollten.

```js
const radioButtons = [...document.querySelectorAll('input[type="radio"]')];
assert(
  radioButtons.every((btn) =>
    btn.getAttribute('name').match(/^indoor-outdoor$/)
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>
          <label><input id="outdoor" type="radio"> Outdoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

