---
id: 5ef9b03c81a63668521804e2
title: Schritt 56
challengeType: 0
dashedName: step-56
---

# --description--

Formulare verwenden häufig Kontrollkästchen für Fragen, bei denen mehr als eine Antwort möglich ist. Das `input`-Element mit einem `type`-Attribut, das auf `checkbox` gesetzt ist, erzeugt ein Kontrollkästchen.

Fügen Sie unter dem gerade hinzugefügten `legend`-Element ein `input` hinzu, dessen `type`-Attribut auf `checkbox` gesetzt ist und dessen Text wie folgt lautet:

` Loving`

# --hints--

Das `input`-Element für Ihr Kontrollkästchen sollte ein öffnendes Tag haben, aber kein schließendes Tag.

```js
assert.isTrue(document.querySelector('fieldset > input') && !code.match(/<\/input\>/g));
```

Sie sollten nur ein input-Element für Ihr Kontrollkästchen hinzugefügt haben. Entfernen Sie alle zusätzlichen.

```js
assert.lengthOf(document.querySelectorAll('fieldset > input'), 1);
```

Ihr neues `input`-Element sollte unter dem `legend`-Element mit dem Text `What's your cat's personality?` stehen. Sie haben sie in der falschen Reihenfolge.

```js
const existingLegendElem = document.querySelectorAll('fieldset > legend')?.[1];
assert.strictEqual(
  existingLegendElem?.nextElementSibling?.nodeName,
  'INPUT'
);
```

Ihr neues `input`-Element hat kein `type`-Attribut. Prüfen Sie, ob nach dem Namen des öffnenden Tags ein Leerzeichen steht.

```js
assert.isTrue(document.querySelector('fieldset > input')?.hasAttribute('type'));
```

Ihr neues `input`-Element sollte ein `type`-Attribut mit dem Wert `checkbox` haben. Sie haben entweder den Wert weggelassen oder einen Tippfehler. Denken Sie daran, dass Attributwerte in Anführungszeichen stehen sollten.

```js
assert.match(document.querySelector('fieldset > input')?.getAttribute('type'),/^checkbox$/i);
```

Obwohl Sie das `input`-Attribut des neuen `type`-Elements auf `checkbox` gesetzt haben, wird empfohlen, den Wert eines Attributs immer in Anführungszeichen zu setzen.

```js
assert(!/\<\s*input\s+type\s*=\s*checkbox/i.test(code));
```

Der Text ` Loving` sollte direkt rechts neben Ihrem Kontrollkästchen stehen. Achten Sie darauf, dass zwischen dem Element und dem Text ein Leerzeichen ist. Sie haben entweder den Text weggelassen oder einen Tippfehler.

```js
const checkboxInputElem = document.querySelector('input[type="checkbox"]');
assert.match(
  checkboxInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '),
  / Loving/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
--fcc-editable-region--
            <legend>What's your cat's personality?</legend>
            
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

