---
id: 5ef9b03c81a63668521804d7
title: Schritt 38
challengeType: 0
dashedName: step-38
---

# --description--

Das Attribut `action` gibt an, wohin die Formulardaten gesendet werden sollen.

Hier ist ein Beispiel für ein `form`-Element mit einem `action`-Attribut:

```html
<form action="/submit-url"></form>
```

Im Beispiel sagt `action="/submit-url"` dem Browser, dass die Formulardaten an den Pfad `/submit-url` gesendet werden sollen.

Fügen Sie dem `action`-Element ein `https://freecatphotoapp.com/submit-cat-photo`-Attribut mit dem Wert `form` hinzu.

# --hints--

Ihr `form`-Element sollte ein öffnendes Tag und ein schließendes Tag in der richtigen Reihenfolge haben. Möglicherweise fehlt Ihnen eines oder beide der erforderlichen Tags, oder sie sind in der falschen Reihenfolge.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

Ihr `form`-Element, das im letzten `section`-Element verschachtelt ist, sollte unter dem `h2`-Element stehen. Sie haben das `h2`-Element und das `form`-Element in der falschen Reihenfolge.

```js
assert.equal(document.querySelector('form')?.previousElementSibling?.nodeName, 'H2');
```

Ihr `form`-Element sollte keinen Inhalt haben. Entfernen Sie alle HTML-Elemente oder Texte zwischen den Tags des `form`-Elements.

```js
assert.lengthOf(document.querySelector('form')?.innerHTML?.trim(), 0);
```

Ihr `form`-Element hat kein `action`-Attribut. Prüfen Sie, ob nach dem Namen des öffnenden Tags ein Leerzeichen steht und/oder vor allen Attributnamen Leerzeichen vorhanden sind.

```js
const form = document.querySelector('form');
assert(form.hasAttribute('action'));
```

Ihr `form`-Element sollte ein `action`-Attribut mit dem Wert `https://freecatphotoapp.com/submit-cat-photo` haben.

```js
const form = document.querySelector('form');
assert(
  form
    .getAttribute('action')
    .match(/^https:\/\/freecatphotoapp\.com\/submit-cat-photo$/)
);
```

Obwohl Sie das `action`-Attribut auf die richtige URL gesetzt haben, wird empfohlen, den Wert eines Attributs immer in Anführungszeichen zu setzen.

```js
assert(
  !/\<form\s+action\s*=\s*https:\/\/freecatphotoapp\.com\/submit-cat-photo/.test(
    code
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form>
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

