---
id: 5ef9b03c81a63668521804d6
title: Schritt 37
challengeType: 0
dashedName: step-37
---

# --description--

Jetzt fügen Sie ein Webformular hinzu, um Informationen von Benutzenden zu sammeln.

Das `form`-Element wird verwendet, um Informationen von einem Benutzer wie Name, E-Mail und andere Details zu erhalten.

Fügen Sie nach der `Cat Form`-Überschrift ein `form`-Element hinzu.

# --hints--

Ihr `form`-Element sollte ein öffnendes Tag und ein schließendes Tag haben. Möglicherweise fehlt Ihnen eines oder beide der erforderlichen Tags, oder sie sind in der falschen Reihenfolge.

```js
assert(document.querySelector('form') && code.match(/<\/form>/g));
```

Die Tags Ihres `form`-Elements sind nicht in der richtigen Reihenfolge.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.indexOf('<form>') < noSpaces.indexOf('</form>'));
```

Das `form`-Element, das im letzten `section`-Element verschachtelt ist, sollte unter dem `h2`-Element stehen. Sie haben das `h2`-Element und das `form`-Element in der falschen Reihenfolge.

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

Das `form`-Element sollte keinen Inhalt haben. Entfernen Sie alle HTML-Elemente oder Texte zwischen den Tags des `form`-Elements.

```js
assert.lengthOf(document.querySelector('form')?.innerHTML?.trim(), 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Form</h2>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

