---
id: 5efb2c990dc218d6c85f89b2
title: Schritt 45
challengeType: 0
dashedName: step-45
---

# --description--

Obwohl Sie Ihre Schaltfläche unterhalb der Texteingabe hinzugefügt haben, erscheinen sie nebeneinander auf der Seite. Das liegt daran, dass sowohl `input` als auch `button` <dfn>Inline-Elemente</dfn> sind, die nicht in neuen Zeilen angezeigt werden.

Die von Ihnen hinzugefügte Schaltfläche sendet das Formular standardmäßig ab. Allerdings kann das Verlassen auf das Standardverhalten zu Verwirrung führen. Fügen Sie dem `type` das `submit`-Attribut mit dem Wert `button` hinzu, um klarzustellen, dass es sich um eine Absenden-Schaltfläche handelt.

# --hints--

Ihr `button`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert(document.querySelector('button'));
```

Ihr `button`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
assert(code.match(/<\/button\>/));
```

Ihr `button`-Element hat kein `type`-Attribut. Prüfen Sie, ob nach dem Namen des öffnenden Tags ein Leerzeichen steht.

```js
assert.isTrue(document.querySelector('button')?.hasAttribute('type'));
```

Ihr `button`-Element sollte ein `type`-Attribut mit dem Wert `submit` haben. Sie haben entweder den Wert weggelassen oder einen Tippfehler. Denken Sie daran, dass Attributwerte in Anführungszeichen stehen sollten.

```js
assert.match(
  document.querySelector('button')?.getAttribute('type'),
  /^submit$/i
);
```

Obwohl Sie das `type`-Attribut des `button`-Elements auf `submit` gesetzt haben, wird empfohlen, den Wert eines Attributs immer in Anführungszeichen zu setzen.

```js
assert(!/\<\s*button\s+type\s*=\s*submit/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button>Submit</button>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

