---
id: 5efb23e70dc218d6c85f89b1
title: Schritt 40
challengeType: 0
dashedName: step-40
---

# --description--

Es gibt viele Arten von Eingaben, die Sie mit dem `type`-Attribut erstellen können. Sie können ganz einfach ein Passwortfeld, eine Zurücksetzen-Schaltfläche oder eine Steuerung erstellen, mit der Benutzer eine Datei von ihrem Computer auswählen können.

Erstellen Sie ein Textfeld, um eine Texteingabe von einem Benutzer zu erhalten, indem Sie dem `type`-Element das `text`-Attribut mit dem Wert `input` hinzufügen.

# --hints--

Sie haben entweder Ihr `input`-Element gelöscht oder es hat eine ungültige Syntax. Wenn Sie Attribute hinzugefügt haben, stellen Sie sicher, dass deren Werte in Anführungszeichen stehen.

```js
assert(document.querySelectorAll('input').length);
```

Ihr `form`-Element sollte nur das `input`-Element enthalten. Entfernen Sie alle zusätzlichen HTML-Elemente oder Texte zwischen den Tags des `form`-Elements.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Ihr `input`-Element hat kein `type`-Attribut mit dem Wert `text`. Überprüfen Sie, ob nach dem Namen des öffnenden Tags ein Leerzeichen steht und/oder ob vor allen Attributnamen Leerzeichen sind.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

Ihr `input`-Element sollte ein `type`-Attribut mit dem Wert `text` haben. Sie haben entweder den Wert ausgelassen oder einen Tippfehler. Denken Sie daran, dass Attributwerte in Anführungszeichen stehen sollten.

```js
assert.match(
  document.querySelector('input')?.getAttribute('type').replace(/\s+/g, ' '),
  /^text$/i
);
```

Obwohl Sie das `type`-Attribut des `input`-Elements auf `text` gesetzt haben, wird empfohlen, den Wert eines Attributs immer in Anführungszeichen zu setzen.

```js
assert(!/\<input\s+type\s*=\s*text/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

