---
id: 5efc4f528d6a74d05e68af74
title: Schritt 58
challengeType: 0
dashedName: step-58
---

# --description--

Es gibt eine weitere Möglichkeit, den Text eines `input`-Elements mit dem Element selbst zu verknüpfen. Sie können den Text innerhalb eines `label`-Elements verschachteln und ein `for`-Attribut mit demselben Wert wie das `input`-Attribut des `id`-Elements hinzufügen.

Gegeben ist ein `input`-Element wie unten:

```html
<input id="breakfast" type="radio" name="meal" value="breakfast">
```

Ein Beispiel für ein `label`-Element, das mit diesem `input`-Element verknüpft ist, lautet:

```html
<label for="breakfast">Breakfast</label>
```

Verknüpfen Sie den Text `Loving` mit dem Kontrollkästchen, indem Sie nur den Text `Loving` in einem `label`-Element verschachteln und ihm ein passendes `for`-Attribut geben.

# --hints--

Sie sollten sicherstellen, dass das Kontrollkästchen-Element weiterhin vorhanden ist und Sie daran keine Änderungen vorgenommen haben. Es sollte so bleiben wie im Ausgangscode.

```js
const checkboxElementRegex = /\s*<input\s+id\s*=\s*"loving"\s*type\s*=\s*"checkbox"\s*>/;
assert.match(code, checkboxElementRegex);
```

Ihr Kontrollkästchen sollte weiterhin ein `id`-Attribut mit dem Wert `loving` haben. Erwartet wurde `--fcc-expected--`, aber gefunden wurde `--fcc-actual--`.

```js
const checkboxElementId = document.querySelector('input[type="checkbox"]')?.getAttribute('id');
assert.strictEqual(checkboxElementId, 'loving');
```

Sie sollten ein neues `label`-Element nach dem Kontrollkästchen-Element hinzufügen. Achten Sie darauf, dass es sowohl ein öffnendes als auch ein schließendes Tag hat.

```js
assert.lengthOf(document.querySelectorAll('label'), 3);
assert.lengthOf(code.match(/<\/label\>/g), 3);
```

Ihr `label`-Element sollte sich nach Ihrem Kontrollkästchen befinden.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElement = checkboxElement?.nextElementSibling?.nodeName;
assert.strictEqual(checkboxSiblingElement, 'LABEL');
```

Ihr `label`-Element sollte den Text `Loving` enthalten.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementText = checkboxElement?.nextElementSibling?.innerText;

assert.strictEqual(checkboxSiblingElementText, 'Loving');
```

Der Bezeichnungstext `Loving` sollte nur innerhalb Ihres neuen `label`-Elements stehen. Achten Sie darauf, den Text nicht außerhalb des `label`-Elements zu duplizieren.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingNode = checkboxElement?.nextSibling?.nodeValue ?? '';
const labelElement = checkboxElement?.nextElementSibling;
const labelSiblingNode = labelElement?.nextSibling?.nodeValue ?? '';

assert.notInclude(checkboxSiblingNode, 'Loving');
assert.notInclude(labelSiblingNode, 'Loving');
```

Ihr `label`-Element hat kein `for`-Attribut. Achten Sie darauf, dass zwischen dem öffnenden `label`-Tag-Namen und dem `for`-Attribut ein Leerzeichen steht.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementAttr = checkboxElement?.nextElementSibling?.getAttribute('for');

assert.isNotNull(checkboxSiblingElementAttr);
```

Ihr `label`-Element sollte ein `for`-Attribut mit dem Wert `loving` haben. Erwartet wurde `--fcc-expected--`, aber gefunden wurde `--fcc-actual--`.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementAttr = checkboxElement?.nextElementSibling?.getAttribute('for');

assert.strictEqual(checkboxSiblingElementAttr, 'loving');
```

Zwischen Ihrem Kontrollkästchen und Ihrem neuen `label`-Element sollte ein Leerzeichen stehen.

```js
assert.match(code, />\s+<label\s+for\s*=\s*('|")loving/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> Loving
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

