---
id: 5f05a1d8e233dff4a68508d8
title: Schritt 49
challengeType: 0
dashedName: step-49
---

# --description--

Erstellen Sie unterhalb der ersten eine weitere Optionsschaltfläche. Verschachteln Sie sie in einem `label`-Element mit `Outdoor` als `label`-Text. Geben Sie der Optionsschaltfläche ein `id`-Attribut mit `outdoor` als Wert.

# --hints--

Nur die ursprüngliche `Indoor`-Optionsschaltfläche `input` sollte ein `id` mit dem Wert `indoor` haben. Sie können den Schritt neu starten, um das ursprüngliche HTML zurückzubekommen, falls nötig.

```js
assert(document.querySelectorAll('#indoor').length < 2);
```

Sie sollten keine Änderungen an der `Indoor`-Optionsschaltfläche vornehmen. Sie können den Schritt neu starten, um das ursprüngliche HTML zurückzubekommen, falls nötig.

```js
const indoorInput = document.querySelectorAll('#indoor');
assert(
  indoorInput.length == 1 &&
  indoorInput[0]?.nodeName?.toUpperCase() === 'INPUT' &&
  indoorInput[0]?.type === 'radio' &&
  code.match(/<label\s*>\s*<input [^>]*(id=["']?indoor["']?)[^>]*>\s*Indoor\s*<\/label>/)
);
```

Sie sollten genau ein neues `input`-Element verschachtelt in einem neuen `label`-Element hinzufügen. Stellen Sie sicher, dass Ihr neues `label` sowohl ein öffnendes als auch ein schließendes Tag hat.

```js
assert(document.querySelectorAll('label input').length === 2);
```

Sie sollten keine Attribute zu einem der öffnenden `label`-Tags hinzufügen.

```js
assert(code.match(/<label\s*>/g).length === 2);
```

Ihr neues `input` sollte ein `id`-Attribut mit dem Wert `outdoor` haben.

```js
assert(document.querySelector('label input[id="outdoor"]'));
```

Ihr neues `input` sollte ein `type`-Attribut mit dem Wert `radio` haben.

```js
assert(document.querySelector('label input[id="outdoor"][type="radio"]'));
```

Sie sollten keine neuen Elemente hinzufügen außer einem `input`, das in einem `label` verschachtelt ist.

```js
assert(document.querySelector('label input[id="outdoor"]:only-child'));
```

Links von Ihrem neuen `input` darf kein Text stehen.

```js
const outdoorBtn = document.querySelector('label input[id="outdoor"]');
assert(
    outdoorBtn?.previousSibling?.nodeName !== '#text' ||
    outdoorBtn?.previousSibling?.nodeValue?.replace(/\s/g, '') === ''
);
```

Der Bezeichnungstext für Ihre neue Optionsschaltfläche muss genau `Outdoor` sein. Das schließt Groß- und Kleinschreibung ein.

```js
assert(document.querySelector('label > input[id="outdoor"]')?.nextSibling?.nodeValue?.replace(/^\s+|\s+$/g, '') === 'Outdoor');
```

Ihre neue Optionsschaltfläche und Bezeichnung sollten unmittelbar unterhalb/nach der `Indoor`-Optionsschaltfläche und Bezeichnung stehen. Es dürfen keine anderen Tags dazwischen sein.

```js
assert(code.match(/<\/label>\s*<label\s*>\s*<input [^>]+>\s*Outdoor/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>
          
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

