---
id: 5f07be6ef7412fbad0c5626b
title: Schritt 18
challengeType: 0
dashedName: step-18
---

# --description--

Bevor Sie neuen Inhalt hinzufügen, sollten Sie ein `section`-Element verwenden, um den Inhalt der Katzenfotos vom zukünftigen Inhalt zu trennen.

Das `section`-Element wird verwendet, um Abschnitte in einem Dokument zu definieren, wie Kapitel, Kopfzeilen, Fußzeilen oder andere Abschnitte des Dokuments. Es ist ein semantisches Element, das bei SEO und Barrierefreiheit hilft.

```html
<section>
  <h2>Section Title</h2>
  <p>Section content...</p>
</section>
```

Nehmen Sie Ihre `h2`-, `p`- und Anker- (`a`) Elemente und verschachteln Sie sie in einem `section`-Element.

# --hints--

Ihr `section`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben folgende Syntax: `<elementName>`.

```js
assert(document.querySelector('section'));
```

Ihr `section`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
assert(code.match(/<\/section\s*>/i));
```

Das gesamte `section`-Element sollte zwischen dem öffnenden und schließenden Tag des `main`-Elements liegen.

```js
assert(document.querySelector('section').parentNode.nodeName === 'MAIN');
```

Das vorhandene `h2`-, zwei `p`-Elemente und das Anker- (`a`) Element sollten zwischen dem öffnenden und schließenden Tag des `section`-Elements liegen.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert.lengthOf(foundElements, 4);
```

Das `h1`-Element sollte nicht im `section`-Element verschachtelt sein.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>

      <h2>Cat Photos</h2>
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

    </main>
--fcc-editable-region--
  </body>
</html>
```

