---
id: 5f07c98cdb9413cbd4b16750
title: Schritt 19
challengeType: 0
dashedName: step-19
---

# --description--

Es ist Zeit, einen neuen Abschnitt hinzuzufügen. Fügen Sie unter dem vorhandenen `section` Element ein zweites `section` Element hinzu.

# --hints--

Ihr `section` Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert(document.querySelectorAll('section').length >= 2);
```

Sie sollten nur ein öffnendes `section` Tag hinzufügen. Bitte entfernen Sie alle zusätzlichen.

```js
assert(document.querySelectorAll('section').length === 2);
```

Ihr `section` Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<` Zeichen.

```js
assert(code.match(/<\/section>/g).length >= 2);
```

Sie sollten nur ein schließendes `section` Tag hinzufügen. Bitte entfernen Sie alle zusätzlichen.

```js
assert(code.match(/<\/section>/g).length === 2);
```

Das zweite `section` Element sollte nicht im ersten `section` Element verschachtelt sein.

```js
const childrenOf1stSection = [
  ...document.querySelector('main > section').children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 0);
```

Beide `section` Elemente sollten sich zwischen dem öffnenden und schließenden Tag des `main` Elements befinden.

```js
const childrenOfMain = [...document.querySelector('main').children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```

