---
id: 5f07fb1579dc934717801375
title: Schritt 35
challengeType: 0
dashedName: step-35
---

# --description--

Es ist Zeit, einen neuen Abschnitt hinzuzufügen. Fügen Sie unter dem zweiten `section`-Element ein drittes `section`-Element hinzu.

# --hints--

Ihr `section`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert(document.querySelectorAll('section').length >= 3);
```

Sie sollten nur ein öffnendes `section`-Tag hinzufügen. Bitte entfernen Sie alle zusätzlichen.

```js
assert(document.querySelectorAll('section').length === 3);
```

Ihr `section`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
assert(code.match(/<\/section>/g).length >= 3);
```

Sie sollten nur ein schließendes `section`-Tag hinzufügen. Bitte entfernen Sie alle zusätzlichen.

```js
assert(code.match(/<\/section>/g).length === 3);
```

Alle `section`-Elemente sollten sich zwischen dem öffnenden und schließenden Tag des `main`-Elements befinden.

```js
const childrenOfMain = [...document.querySelector('main').children];
const sectionElemsFound = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(sectionElemsFound.length === 3);
```

Das letzte `section`-Element sollte keinen Inhalt haben. Entfernen Sie alle HTML-Elemente oder Texte innerhalb des `section`-Elements.

```js
assert.lengthOf(document.querySelectorAll('main > section')?.[2]?.children, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```

