---
id: 5f1a80975fc4bcae0edb3497
title: Schritt 51
challengeType: 0
dashedName: step-51
---

# --description--

Wenn Sie die Optionsschaltfläche `Indoor` auswählen und das Formular absenden, basieren die Formulardaten für die Schaltfläche auf ihren `name`- und `value`-Attributen. Da Ihre Optionsschaltflächen kein `value`-Attribut haben, enthalten die Formulardaten `indoor-outdoor=on`, was nicht nützlich ist, wenn Sie mehrere Schaltflächen haben.

Fügen Sie beiden Optionsschaltflächen ein `value`-Attribut hinzu. Für mehr Komfort setzen Sie das `value`-Attribut der Schaltfläche auf denselben Wert wie ihr `id`-Attribut.

# --hints--

Beide Optionsschaltflächen sollten sich weiterhin zwischen öffnenden und schließenden `label`-Element-Tags befinden.

```js
const labelChildNodes = [...document.querySelectorAll('form > label')].map(
  (node) => node.childNodes
);
assert(
  labelChildNodes.filter((childNode) => childNode[0].nodeName === 'INPUT')
    .length === 2
);
```

Beide Optionsschaltflächen sollten ein `value`-Attribut haben. Prüfen Sie, ob nach dem Namen des öffnenden Tags ein Leerzeichen ist und/oder vor allen Attributnamen Leerzeichen stehen.

```js
const radioButtons = [...document.querySelectorAll('input[type="radio"]')];
assert(radioButtons.every((btn) => btn.hasAttribute('value')));
```

Beide Optionsschaltflächen sollten ein `type`-Attribut haben. Prüfen Sie, ob nach dem Namen des öffnenden Tags ein Leerzeichen ist und/oder vor allen Attributnamen Leerzeichen stehen.

```js
const radioButtons = [...document.querySelectorAll('input')];
assert(radioButtons.every((btn) => btn.hasAttribute('type')));
```

Das `Indoor`-Optionsschaltflächen-`value`-Attribut sollte auf `indoor` gesetzt sein. Sie haben entweder den Wert ausgelassen oder einen Tippfehler. Denken Sie daran, dass Attributwerte in Anführungszeichen stehen sollten.

```js
const indoorRadioButton = document.querySelector('#indoor');
assert(indoorRadioButton.getAttribute('value').match(/^indoor$/));
```

Das `Outdoor`-Optionsschaltflächen-`value`-Attribut sollte auf `outdoor` gesetzt sein. Sie haben entweder den Wert ausgelassen oder einen Tippfehler. Denken Sie daran, dass Attributwerte in Anführungszeichen stehen sollten.

```js
const outdoorRadioButton = document.querySelector('#outdoor');
assert(outdoorRadioButton.getAttribute('value').match(/^outdoor$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
          <label><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

