---
id: 7cf9b03d81a65668421804c3
title: Schritt 41
challengeType: 0
dashedName: step-41
---

# --description--

Damit die Daten eines Formulars an den im `action`-Attribut angegebenen Ort zugreifen können, müssen Sie dem Textfeld ein `name`-Attribut zuweisen und ihm einen Wert geben, der die zu übermittelnden Daten darstellt.

Hier ist ein Beispiel für ein `input`-Element mit einem `name`-Attribut:

```html
<input type="text" name="name">
```

Fügen Sie dem Textfeld das `name`-Attribut mit dem Wert `catphotourl` hinzu.

# --hints--

Sie haben entweder Ihr `input`-Element gelöscht oder es hat eine ungültige Syntax. Alle Attributwerte sollten in Anführungszeichen stehen.

```js
assert(document.querySelectorAll('input').length);
```

Ihr `form` sollte nur das `input`-Element enthalten. Entfernen Sie alle zusätzlichen HTML-Elemente oder Texte innerhalb des `form`-Elements.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
  document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Ihr `input`-Element hat kein `name`-Attribut. Prüfen Sie, ob nach dem öffnenden Tag-Namen ein Leerzeichen steht und/oder vor allen Attributnamen Leerzeichen vorhanden sind.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('name'));
```

Ihr `input`-Element sollte ein `name`-Attribut mit dem Wert `catphotourl` haben. Sie haben entweder den Wert weggelassen oder einen Tippfehler. Denken Sie daran, dass Attributwerte in Anführungszeichen stehen sollten.

```js
assert.match(
  document.querySelector('input')?.getAttribute('name'),
  /^catphotourl$/i
);
```

Obwohl Sie das `name`-Attribut des `input`-Elements auf `catphotourl` gesetzt haben, wird empfohlen, den Wert eines Attributs immer in Anführungszeichen zu setzen.

```js
assert(!/\<\s*input\s+.*\s*=\s*catphotourl/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

