---
id: 696920c0c98a1ed58eb86293
title: 商品ショーケースを作成する
challengeType: 25
dashedName: lab-product-showcase
demoType: onClick
---

# --description--

このラボでは、すべてのHTMLとCSSが提供されています。TypeScriptを使って商品ショーケースを完成させてください。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了させてください。

**ユーザーストーリー:**

1. `Item`という名前のインターフェースを作成してください：
   - `type`プロパティは`"book"`、`"electronics"`、または`"clothing"`のいずれかのみを許容します。
   - `id`プロパティは`string`型です。
   - `price`プロパティは`number`型です。

1. `Book`という名前のインターフェースを作成し、`Item`を拡張してください。`type`プロパティはリテラル`"book"`に設定します。また、`title`と`author`プロパティは`string`型です。

1. `Electronics`という名前のインターフェースを作成し、`Item`を拡張してください。`type`プロパティはリテラル`"electronics"`に設定します。また、`item`と`model`プロパティは`string`型で、オプションの`warranty`プロパティは`number`型です。

1. `Clothing`という名前のインターフェースを作成し、`Item`を拡張してください。`type`プロパティはリテラル`"clothing"`に設定します。また、`item`と`brand`プロパティは`string`型で、オプションの`size`プロパティは`"S"`、`"M"`、または`"L"`のいずれかのみを許容します。

1. `Product`という名前のユニオン型を作成し、`Book`、`Electronics`、および`Clothing`インターフェースを組み合わせてください。

1. `Collection`という名前のジェネリッククラスを作成し、型パラメータ`T`を持たせてください。コンストラクタは`T`型の配列を受け取り、`items`という名前のプロパティに格納します。

1. `Collection`クラスには、コレクション内のすべてのアイテムを返す`getAll()`メソッドが必要です。

1. `Collection`クラスには、コールバックを受け取り、`filter()`配列の要素のうち条件を満たすものだけを返す`items`メソッドが必要です。コールバックは`T`型の要素を受け取り、ブール値を返します。

1. `renderProduct`という名前の関数を作成してください。パラメータは`Product`型です。次の内容を含む商品を表すHTML文字列を返します：
   - クラス`item`を持ち、IDが商品の`id`の値である要素。
   - クラス`price`を持ち、商品の`price`のテキストを含む要素。
   - 追加の商品情報を含む要素。

1. `renderProduct`関数は型絞り込みを使って、商品タイプに応じて異なるHTMLをレンダリングしてください：

   - 書籍の場合、レンダリングされるHTMLは`Book: [title] by [author]`の形式で、`[title]`と`[author]`はそれぞれタイトルと著者に置き換えます。
   - 電子機器の場合、レンダリングされるHTMLは`Electronics: [item] - [model]`の形式です。保証がある場合は` - Warranty: [warranty] year(s)`を追加します。`[item]`、`[model]`、`[warranty]`はそれぞれアイテム、モデル、保証に置き換えます。
   - 衣類の場合、レンダリングされるHTMLは`Clothing: [item] by [brand]`の形式です。サイズがある場合は` - Size [size]`を追加します。`[item]`、`[brand]`、`[size]`はそれぞれアイテム、ブランド、サイズに置き換えます。
   - 無効な商品タイプがあった場合は、次の形式のエラーをスローしてください：`Unknown product type: [product]`。`[product]`は無効な商品オブジェクトに置き換えます。無効な商品を文字列に変換するには`JSON.stringify`を使います。

1. 各商品タイプのアイテムを少なくとも1つ含む`Collection<Product>`インスタンスを作成し、`products`という変数に代入してください。

1. `showProducts`という関数を作成し、`#output`要素に商品を表示してください。オプションのフィルターとして商品タイプを受け取ります。省略した場合はすべての商品を表示し、指定した場合はそのタイプの商品だけを表示します。

1. `renderProduct`関数を使って各商品をHTML文字列に変換し、結合した結果を`#output`要素に挿入して商品をレンダリングしてください。

1. `#all`、`#books`、`#electronics`、`#clothing`ボタンをクリックすると、それぞれ対応する商品セットが`#output`要素に表示されます。

1. ページ読み込み時には、デフォルトで全商品が表示されます。

# --hints--

`Item`という名前のインターフェースを作成してください。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Item);
```

`Item`インターフェースには、`type`、`"book"`、または`"electronics"`のみを許容する`"clothing"`プロパティが必要です。

```js
const explorer = await __helpers.Explorer(code);
const { Item } = explorer.interfaces;
const typeProp = Item.typeProps.type;
assert.isTrue(typeProp.isUnionOf(["'book'", "'electronics'", "'clothing'"]));
```

`Item`インターフェースには`id`型の`string`プロパティが必要です。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Item.hasTypeProps({name: "id", type: "string"}));
```

`Item`インターフェースには`price`型の`number`プロパティが必要です。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Item.hasTypeProps({name: "price", type: "number"}));
```

`Book`という名前のインターフェースを作成してください。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Book);
```

`Book`インターフェースは`Item`を拡張してください。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.doesExtend("Item"));
```

`Book`インターフェースには`type`に設定された`"book"`プロパティが必要です。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "type", type: "'book'"}));
```

`Book`インターフェースには`title`型の`string`プロパティが必要です。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "title", type: "string"}));
```

`Book`インターフェースには`author`型の`string`プロパティが必要です。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "author", type: "string"}));
```

`Electronics`という名前のインターフェースを作成してください。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Electronics);
```

`Electronics`インターフェースは`Item`を拡張してください。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.doesExtend("Item"));
```

`Electronics`インターフェースには`type`に設定された`"electronics"`プロパティが必要です。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "type", type: "'electronics'"}));
```

`Electronics`インターフェースには`item`型の`string`プロパティが必要です。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "item", type: "string"}));
```

`Electronics`インターフェースには`model`型の`string`プロパティが必要です。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "model", type: "string"}));
```

`Electronics`インターフェースにはオプションの`warranty`プロパティがあり、`number`型です。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "warranty", type: "number", isOptional: true}));
```

`Clothing`という名前のインターフェースを作成してください。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Clothing);
```

`Clothing`インターフェースは`Item`を拡張してください。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.doesExtend("Item"));
```

`Clothing`インターフェースには`type`に設定された`"clothing"`プロパティが必要です。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "type", type: "'clothing'"}));
```

`Clothing`インターフェースには`item`型の`string`プロパティが必要です。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "item", type: "string"}));
```

`Clothing`インターフェースには`brand`型の`string`プロパティが必要です。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "brand", type: "string"}));
```

`Clothing`インターフェースにはオプションの`size`プロパティがあり、`"S"`、`"M"`、または`"L"`のいずれかのみを許容します。

```js
const explorer = await __helpers.Explorer(code);
const { Clothing } = explorer.interfaces;
assert.isTrue(Clothing.hasTypeProps({name: "size", isOptional: true}));
assert.isTrue(Clothing.typeProps.size.isUnionOf(["'S'", "'M'", "'L'"]));
```

`Product`という名前のユニオン型を作成し、`Book`、`Electronics`、および`Clothing`インターフェースを組み合わせてください。

```js
const explorer = await __helpers.Explorer(code);
const { Product } = explorer.types;
assert.isTrue(Product.isUnionOf(["Book", "Electronics", "Clothing"]));
```

`Collection`クラスを作成し、型パラメータ`T`を持たせてください。

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
assert.isTrue(Collection.typeParameters[0].matches("T"));
```

`Collection`クラスのコンストラクタは`T`型の配列を受け取り、`items`という名前のプロパティに格納します。

```js
let arr = [1, 2, 3];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.items, arr);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.items, arr);
```

`Collection`クラスには`getAll`メソッドが必要です。

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
console.log(Collection.methods)
console.log(Collection.methods.getAll)
assert.exists(Collection.methods.getAll);
```

`getAll`メソッドはコレクション内のすべてのアイテムを返します。

```js
let arr = [1, 2, 3];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.getAll(), arr);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.getAll(), arr);
```

`Collection`クラスには`filter`メソッドが必要です。

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
assert.exists(Collection.methods.filter);
```

`filter`メソッドはコールバック関数を受け取り、`items`プロパティの要素のうち条件を満たすものを含む配列を返します。

```js
let arr = [1, 2, 3, 4, 5, 6];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.filter(n => n % 2 === 0), [2, 4, 6]);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.filter(s => s === "c"), ["c"]);
```

`renderProduct`という名前の関数を作成してください。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.renderProduct);
```

`renderProduct(product)`はクラス`item`を持ち、IDが`product.id`と等しい要素を含む文字列を返します。

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  assert.isNotNull(testEl.querySelector(`[class="item"][id="${p.id}"]`));
})
```

`renderProduct(product)`はクラス`price`を持ち、テキストが`product.price`と等しい要素を含む文字列を返します。

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const price = testEl.querySelector(`[class="price"]`);
  assert.include(price.innerText, p.price.toString());
})
```

`renderProduct`が`Book`型の商品`p`で呼び出された場合、`Book: [title] by [author]`の形式に従ったテキストを含む要素を返します。`[title]`と`[author]`は`p`の対応するプロパティに置き換えます。

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
  { id: "b4", type: "book", title: "The Lord of the Rings", author: "J.R.R. Tolkien", price: 24.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const expectedText = `Book: ${p.title} by ${p.author}`;
  assert.include(testEl.innerText, expectedText);
})
```

`renderProduct`が`Electronics`型の商品`p`で呼び出された場合、`Electronics: [item] - [model]`の形式か、`Electronics: [item] - [model] - Warranty: [warranty] year(s)`が未定義でなければ`p.warranty`の形式に従ったテキストを含む要素を返します。`[item]`、`[model]`、`[warranty]`は`p`の対応するプロパティに置き換えます。

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "e2", type: "electronics", item: "Smartphone", model: "NovaCore X1-Alpha", warranty: 2, price: 699.99 },
  { id: "e3", type: "electronics", item: "Headphones", model: "EchoSphere Silent-7", price: 129.99 },
  { id: "e4", type: "electronics", item: "Laptop", model: "HexaBook Orion-15", warranty: 3, price: 1199.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const baseFormat = `Electronics: ${p.item} - ${p.model}`;
  const expectedText = !p.warranty ? baseFormat : baseFormat + ` - Warranty: ${p.warranty} year(s)`;
  assert.include(testEl.innerText, expectedText);
  if (!p.warranty) {
    assert.notInclude(testEl.innerText, `undefined`);
  }
})
```

`renderProduct`が`Clothing`型の商品`p`で呼び出された場合、`Clothing: [item] by [brand]`の形式か、`Clothing: [item] by [brand] - Size [size]`が未定義でなければ`p.size`の形式に従ったテキストを含む要素を返します。`[item]`、`[brand]`、`[size]`は`p`の対応するプロパティに置き換えます。

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c4", type: "clothing", item: "T-Shirt", brand: "Fabricon", price: 19.99 },
  { id: "c3", type: "clothing", item: "Jeans", brand: "BlueWeave", size: "L", price: 59.99 },
  { id: "c2", type: "clothing", item: "Hoodie", brand: "CozyForge", size: "S", price: 49.99 },
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const baseFormat = `Clothing: ${p.item} by ${p.brand}`;
  const expectedText = !p.size ? baseFormat : baseFormat + ` - Size ${p.size}`;
  assert.include(testEl.innerText, expectedText);
  if (!p.size) {
    assert.notInclude(testEl.innerText, `undefined`);
  }
})
```

`renderProduct`が無効な商品で呼び出された場合、`Unknown product type: [p]`の形式のエラーをスローします。`[p]`は`JSON.stringify`で文字列化された無効な商品に置き換えます。

```js
const invalidProduct = { item: "T-Shirt", price: 19.99 };
const expected = `Unknown product type: ${JSON.stringify(invalidProduct)}`;
try {
  renderProduct(invalidProduct);
  assert.fail();
} catch({name, message}) {
  if (name == "AssertionError") {
    assert.fail("renderProduct should throw with invalid product");
  } else {
    assert.equal(message, expected);
  }
}
```

`Collection<Product>`インスタンスを作成し、`products`という変数に代入してください。

```js
assert.instanceOf(products, Collection);
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.products.value.typeArguments[0].matches("Product"));
```

`products`変数には各商品タイプのアイテムが少なくとも1つ含まれている必要があります。

```js
assert.isTrue(products.items.some(p => p.type === "book"));
assert.isTrue(products.items.some(p => p.type === "electronics"));
assert.isTrue(products.items.some(p => p.type === "clothing"));
```

`showProducts`という名前の関数を作成してください。

```js
assert.isFunction(showProducts);
```

`#all`ボタンをクリックすると、`#output`要素にすべての商品が表示されます。

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
books.dispatchEvent(new Event("click"));
all.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), len);
```

`#books`ボタンをクリックすると、対応する商品セットが`#output`要素に表示されます。

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "book") {
    count++;
  }
});
clothing.dispatchEvent(new Event("click"));
books.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

`#electronics`ボタンをクリックすると、対応する商品セットが`#output`要素に表示されます。

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "electronics") {
    count++;
  }
});
clothing.dispatchEvent(new Event("click"));
electronics.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

`#clothing`ボタンをクリックすると、対応する商品セットが`#output`要素に表示されます。

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "clothing") {
    count++;
  }
});
books.dispatchEvent(new Event("click"));
clothing.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

ページ読み込み時には、デフォルトで全商品が表示されます。これを保証するには`DOMContentLoaded`イベントを使ってください。

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
document.dispatchEvent(new Event("DOMContentLoaded"));
assert.lengthOf(output.querySelectorAll(".item"), len);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Product Showcase</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Product Showcase</h1>
  <div class="buttons">
    <button id="all">All</button>
    <button id="books">Books</button>
    <button id="electronics">Electronics</button>
    <button id="clothing">Clothing</button>
  </div>
  <div id="output" class="product-list"></div>
  <script src="index.ts"></script>
</body>
</html>
```

```css
body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  background: linear-gradient(135deg, #0f4c75 0%, #3282b8 100%);
  padding: 40px 20px;
  margin: 0;
  color: #1a1a1a;
  min-height: 100vh;
}

h1 {
  text-align: center;
  margin-bottom: 40px;
  color: #ffffff;
  font-size: 2.5rem;
  font-weight: 600;
  letter-spacing: -0.5px;
}

.buttons {
  display: flex;
  justify-content: center;
  gap: 12px;
  margin-bottom: 40px;
  flex-wrap: wrap;
}

button {
  padding: 12px 24px;
  font-size: 15px;
  font-weight: 500;
  cursor: pointer;
  border: none;
  border-radius: 8px;
  background: rgba(255, 255, 255, 0.9);
  color: #333;
  transition: all 0.3s ease-in-out;
  box-shadow: 0 2px 8px rgba(0, 0, 0, 0.1);
}

button:hover {
  background: white;
  transform: translateY(-2px);
  box-shadow: 0 4px 12px rgba(0, 0, 0, 0.15);
}

button.active {
  background: #ff9800;
  color: #1a1a1a;
  box-shadow: 0 4px 15px rgba(255, 152, 0, 0.4);
  font-weight: 600;
}

#output {
  display: grid;
  grid-template-columns: repeat(auto-fill, minmax(300px, 1fr));
  gap: 24px;
  max-width: 1200px;
  margin: 0 auto;
}

.item {
  background: white;
  border-radius: 12px;
  padding: 24px;
  box-shadow: 0 4px 15px rgba(0, 0, 0, 0.1);
  transition: all 0.3s ease-in-out;
  display: flex;
  flex-direction: column;
  gap: 12px;
}

.item:hover {
  transform: translateY(-8px);
  box-shadow: 0 12px 25px rgba(0, 0, 0, 0.15);
}

.item strong {
  font-size: 12px;
  text-transform: uppercase;
  letter-spacing: 1px;
  color: #0f4c75;
  font-weight: 700;
}

.item > div:not(.price) {
  font-size: 16px;
  color: #555;
  line-height: 1.5;
}

.price {
  margin-top: auto;
  font-size: 24px;
  font-weight: 700;
  color: #0f4c75;
  padding-top: 12px;
  border-top: 2px solid #f0f0f0;
}

@media (max-width: 768px) {
  #output {
    grid-template-columns: repeat(auto-fill, minmax(250px, 1fr));
    gap: 16px;
  }

  h1 {
    font-size: 2rem;
  }

  .item {
    padding: 16px;
  }
}

@media (max-width: 480px) {
  #output {
    grid-template-columns: 1fr;
  }

  button {
    padding: 10px 16px;
    font-size: 14px;
  }

  .price {
    font-size: 20px;
  }
}
```

```ts

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Product Showcase</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Product Showcase</h1>
  <div class="buttons">
    <button id="all">All</button>
    <button id="books">Books</button>
    <button id="electronics">Electronics</button>
    <button id="clothing">Clothing</button>
  </div>
  <div id="output" class="product-list"></div>
  <script src="index.ts"></script>
</body>
</html>
```

```css
body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  background: linear-gradient(135deg, #0f4c75 0%, #3282b8 100%);
  padding: 40px 20px;
  margin: 0;
  color: #1a1a1a;
  min-height: 100vh;
}

h1 {
  text-align: center;
  margin-bottom: 40px;
  color: #ffffff;
  font-size: 2.5rem;
  font-weight: 600;
  letter-spacing: -0.5px;
}

.buttons {
  display: flex;
  justify-content: center;
  gap: 12px;
  margin-bottom: 40px;
  flex-wrap: wrap;
}

button {
  padding: 12px 24px;
  font-size: 15px;
  font-weight: 500;
  cursor: pointer;
  border: none;
  border-radius: 8px;
  background: rgba(255, 255, 255, 0.9);
  color: #333;
  transition: all 0.3s ease-in-out;
  box-shadow: 0 2px 8px rgba(0, 0, 0, 0.1);
}

button:hover {
  background: white;
  transform: translateY(-2px);
  box-shadow: 0 4px 12px rgba(0, 0, 0, 0.15);
}

button.active {
  background: #ff9800;
  color: #1a1a1a;
  box-shadow: 0 4px 15px rgba(255, 152, 0, 0.4);
  font-weight: 600;
}

#output {
  display: grid;
  grid-template-columns: repeat(auto-fill, minmax(300px, 1fr));
  gap: 24px;
  max-width: 1200px;
  margin: 0 auto;
}

.item {
  background: white;
  border-radius: 12px;
  padding: 24px;
  box-shadow: 0 4px 15px rgba(0, 0, 0, 0.1);
  transition: all 0.3s ease-in-out;
  display: flex;
  flex-direction: column;
  gap: 12px;
}

.item:hover {
  transform: translateY(-8px);
  box-shadow: 0 12px 25px rgba(0, 0, 0, 0.15);
}

.item strong {
  font-size: 12px;
  text-transform: uppercase;
  letter-spacing: 1px;
  color: #0f4c75;
  font-weight: 700;
}

.item > div:not(.price) {
  font-size: 16px;
  color: #555;
  line-height: 1.5;
}

.price {
  margin-top: auto;
  font-size: 24px;
  font-weight: 700;
  color: #0f4c75;
  padding-top: 12px;
  border-top: 2px solid #f0f0f0;
}

@media (max-width: 768px) {
  #output {
    grid-template-columns: repeat(auto-fill, minmax(250px, 1fr));
    gap: 16px;
  }

  h1 {
    font-size: 2rem;
  }

  .item {
    padding: 16px;
  }
}

@media (max-width: 480px) {
  #output {
    grid-template-columns: 1fr;
  }

  button {
    padding: 10px 16px;
    font-size: 14px;
  }

  .price {
    font-size: 20px;
  }
}
```

```ts
interface Item {
  type: "book" | "electronics" | "clothing";
  id: string;
  price: number;
}

interface Book extends Item {
  type: "book";
  title: string;
  author: string;
}

interface Electronics extends Item {
  type: "electronics";
  item: string;
  model: string;
  warranty?: number;
}

interface Clothing extends Item {
  type: "clothing";
  item: string;
  brand: string;
  size?: "S" | "M" | "L";
}

type Product = Book | Electronics | Clothing;

class Collection<T> {
  items: T[];
  constructor(items: T[]) {
    this.items = items;
  }

  getAll(): T[] {
    return this.items;
  }

  filter(callback: (item: T) => boolean): T[] {
    return this.items.filter(callback);
  }
}

function renderProduct(p: Product): string {
  const createProductCard = (id: string, content: string, price: number): string => {
    return `
     <div class="item" id="${id}">
       ${content}
       <div class="price">$${price}</div>
     </div>
   `;
  };

  if (p.type === "book") {
    const content = `<strong>Book:</strong> ${p.title} by ${p.author}`;
    return createProductCard(p.id, content, p.price);
  }

  if (p.type === "electronics") {
    const warranty = p.warranty ? ` - Warranty: ${p.warranty} year(s)` : "";
    const content = `<strong>Electronics:</strong> ${p.item} - ${p.model}${warranty}`;
    return createProductCard(p.id, content, p.price);
  }

  if (p.type === "clothing") {
    const size = p.size ? ` - Size ${p.size}` : "";
    const content = `<strong>Clothing:</strong> ${p.item} by ${p.brand}${size}`;
    return createProductCard(p.id, content, p.price);
  }

  const _never: never = p;
  throw new Error(`Unknown product type: ${JSON.stringify(p)}`);
}


const products = new Collection<Product>([
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },

  { id: "e2", type: "electronics", item: "Smartphone", model: "NovaCore X1-Alpha", warranty: 2, price: 699.99 },
  { id: "c2", type: "clothing", item: "Hoodie", brand: "CozyForge", size: "S", price: 49.99 },
  { id: "b2", type: "book", title: "1984", author: "George Orwell", price: 9.99 },

  { id: "e3", type: "electronics", item: "Headphones", model: "EchoSphere Silent-7", price: 129.99 },
  { id: "c3", type: "clothing", item: "Jeans", brand: "BlueWeave", size: "L", price: 59.99 },
  { id: "e4", type: "electronics", item: "Laptop", model: "HexaBook Orion-15", warranty: 3, price: 1199.99 },

  { id: "b3", type: "book", title: "Brave New World", author: "Aldous Huxley", price: 11.99 },
  { id: "c4", type: "clothing", item: "T-Shirt", brand: "Fabricon", price: 19.99 },
  { id: "e5", type: "electronics", item: "Smartwatch", model: "Chronex Pulse-Q", warranty: 1, price: 199.99 },

  { id: "b4", type: "book", title: "The Lord of the Rings", author: "J.R.R. Tolkien", price: 24.99 },
]);

const output = document.getElementById("output");

// Track currently active filter category (or undefined for all)
let currentFilter: Product["type"] | undefined = undefined;

function showProducts(filter?: Product["type"]): void {
  if (!output) return;

  let itemsToShow: Product[];

  // Toggle logic
  if (currentFilter === filter || filter === undefined) {
    // Reset to all if same button clicked again or filter is undefined
    itemsToShow = products.getAll();
    currentFilter = undefined;
    setActiveButton(undefined);
  } else {
    itemsToShow = products.filter((p: Product): boolean => p.type === filter);
    currentFilter = filter;
    setActiveButton(filter);
  }

  output.innerHTML = itemsToShow.map(renderProduct).join("");
}

// Active Button Highlighting
function setActiveButton(filter: Product["type"] | undefined): void {
  const buttons: { [key: string]: HTMLElement | null } = {
    all: document.getElementById("all"),
    book: document.getElementById("books"),
    electronics: document.getElementById("electronics"),
    clothing: document.getElementById("clothing")
  };

  for (const key in buttons) {
    buttons[key]?.classList.remove("active");
  }

  if (filter === undefined) {
    buttons["all"]?.classList.add("active");
  } else {
    buttons[filter]?.classList.add("active");
  }
}

// Button Event Handlers
document.getElementById("all")?.addEventListener("click", (): void => showProducts());
document.getElementById("books")?.addEventListener("click", (): void => showProducts("book"));
document.getElementById("electronics")?.addEventListener("click", (): void => showProducts("electronics"));
document.getElementById("clothing")?.addEventListener("click", (): void => showProducts("clothing"));

// Initial Render
document.addEventListener("DOMContentLoaded", (): void => showProducts());
```
