---
id: 5d822fd413a79914d39e98cb
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Далі додайте початкові й кінцеві теги `head` та `body` до елемента `html`.

# --hints--

Ви повинні мати початковий тег `head`.

```js
assert(code.match(/<head\s*>/i));
```

Ви повинні мати кінцевий тег `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Ви повинні мати початковий теґ `body`.

```js
assert(code.match(/<body\s*>/i));
```

Ви повинні мати кінцевий теґ `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Елементи `head` і `body` повинні бути братськими.

```js
assert(document.querySelector('head').nextElementSibling.localName === 'body');
```

Елемент `head` повинен бути в межах елемента `html`.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'head'));
```

Елемент `body` повинен бути в межах елемента `html`.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html lang="en">

</html>
--fcc-editable-region--

```

