---
id: 587d7fb7367417b2b2512c0b
title: Використання model.find() для пошуку у вашій базі даних
challengeType: 2
forumTopicId: 301543
dashedName: use-model-find-to-search-your-database
---

# --description--

У найпростішому використанні `Model.find()` приймає документ запиту (JSON-об’єкт) як перший аргумент, а потім функцію зворотного виклику. Він повертає масив співпадінь. Підтримує надзвичайно широкий спектр опцій пошуку. Детальніше читайте в документації.

# --instructions--

Змініть функцію `findPeopleByName`, щоб знайти всіх людей із заданим іменем, використовуючи <code>Model.find() -\> [Person]</code>

Використайте аргумент функції `personName` як ключ пошуку.

# --hints--

Знайти всі елементи, що відповідають критеріям, має вийти успішно

```js
  const response = await fetch(code + '/_api/find-all-by-name', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'r@nd0mN4m3',
      age: 24,
      favoriteFoods: ['pizza']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an Array');
  assert.equal(
    data[0].name,
    'r@nd0mN4m3',
    'item.name is not what expected'
  );
  assert.equal(data[0].__v, 0, 'The item should be not previously edited');
```
