---
id: 587d7fb8367417b2b2512c10
title: Видалити один документ за допомогою model.findByIdAndRemove
challengeType: 2
forumTopicId: 301539
dashedName: delete-one-document-using-model-findbyidandremove
---

# --description--

`findByIdAndRemove` і `findOneAndRemove` схожі на попередні методи оновлення. Вони передають видалений документ у базу даних. Як завжди, використовуйте аргумент функції `personId` як ключ пошуку.

# --instructions--

Змініть функцію `removeById`, щоб видалити одну особу за `_id` цієї особи. Ви повинні використати один із методів `findByIdAndRemove()` або `findOneAndRemove()`.

# --hints--

Видалення елемента має бути успішним

```js
  const response = await fetch(code + '/_api/remove-one-person', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Jason Bourne',
      age: 36,
      favoriteFoods: ['apples']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Jason Bourne', 'item.name is not what expected');
  assert.equal(data.age, 36, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['apples'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0);
  assert.equal(data.count, 0, 'the db items count is not what expected');
```
