---
id: 587d7fb8367417b2b2512c0f
title: Виконати нові оновлення документа за допомогою model.findOneAndUpdate()
challengeType: 2
forumTopicId: 301542
dashedName: perform-new-updates-on-a-document-using-model-findoneandupdate
---

# --description--

Останні версії Mongoose мають методи, які спрощують оновлення документів. Деякі більш просунуті функції (наприклад, pre/post хуки, валідація) поводяться інакше з таким підходом, тому класичний метод досі корисний у багатьох ситуаціях. `findByIdAndUpdate()` можна використовувати для пошуку за ідентифікатором.

# --instructions--

Змініть функцію `findAndUpdate`, щоб знайти людину за `Name` і встановити вік людини на `20`. Використайте параметр функції `personName` як ключ пошуку.

**Примітка:** Ви повинні повернути оновлений документ. Для цього потрібно передати об’єкт опцій `{ new: true }` як третій аргумент у `findOneAndUpdate()`. За замовчуванням ці методи повертають необроблений об’єкт.

# --hints--

findOneAndUpdate елемента має виконуватися успішно

```js
  const response = await fetch(code + '/_api/find-one-update', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Dorian Gray',
      age: 35,
      favoriteFoods: ['unknown']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Dorian Gray', 'item.name is not what expected');
  assert.equal(data.age, 20, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['unknown'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(
    data.__v,
    0,
    'findOneAndUpdate does not increment version by design!'
  );
```
