---
id: 587d7fb6367417b2b2512c09
title: Створення та збереження запису моделі
challengeType: 2
forumTopicId: 301536
dashedName: create-and-save-a-record-of-a-model
---

# --description--

У цьому завданні вам потрібно створити та зберегти запис моделі.

# --instructions--

Усередині функції `createAndSavePerson` створіть екземпляр документа за допомогою конструктора моделі `Person`, який ви створили раніше. Передайте конструктору об’єкт із полями `name`, `age` та `favoriteFoods`. Їхні типи мають відповідати типам у `personSchema`. Потім викличте метод `document.save()` на отриманому екземплярі документа. Передайте йому функцію зворотного виклику за конвенцією Node. Це загальний шаблон; усі наступні методи CRUD приймають функцію зворотного виклику як останній аргумент.

```js
/* Example */

// ...
person.save(function(err, data) {
  //   ...do your stuff here...
});
```

# --hints--

Створення та збереження елемента бази даних має пройти успішно

```js
  const response = await fetch(code + '/_api/create-and-save-person');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isString(data.name, '"item.name" should be a String');
  assert.isNumber(data.age, '28', '"item.age" should be a Number');
  assert.isArray(
    data.favoriteFoods,
    '"item.favoriteFoods" should be an Array'
  );
  assert.equal(data.__v, 0, 'The db item should be not previously edited');
```
