---
id: 587d7fb8367417b2b2512c0e
title: 'Виконання класичних оновлень через пошук, редагування та збереження'
challengeType: 2
forumTopicId: 301541
dashedName: perform-classic-updates-by-running-find-edit-then-save
---

# --description--

Колись це було те, що потрібно було зробити, якщо ви хотіли відредагувати документ і мати змогу якось його використати (наприклад, надіслати назад у відповіді сервера). У Mongoose є спеціальний метод оновлення: `Model.update()`. Він пов’язаний із низькорівневим драйвером mongo. Він може масово редагувати багато документів, що відповідають певним критеріям, але не повертає оновлений документ, а лише повідомлення про «статус». Крім того, він ускладнює валідацію моделей, бо просто напряму викликає драйвер mongo.

# --instructions--

Змініть функцію `findEditThenSave`, щоб знайти людину за `_id` (використайте будь-який із наведених вище методів) з параметром `personId` як ключем пошуку. Додайте `"hamburger"` до списку `favoriteFoods` цієї людини (можна використати `Array.push()`). Потім — всередині функції зворотного виклику пошуку — `save()` оновлений `Person`.

**Примітка:** Це може бути складно, якщо у вашій схемі ви оголосили `favoriteFoods` як масив, не вказуючи тип (тобто `[String]`). У такому разі `favoriteFoods` за замовчуванням має тип Mixed, і вам доведеться вручну позначити його як відредагований, використовуючи `document.markModified('edited-field')`. Детальніше дивіться в нашій <a href="https://www.freecodecamp.org/news/introduction-to-mongoose-for-mongodb-d2a7aa593c57/" target="_blank" rel="noopener noreferrer nofollow">статті про Mongoose</a>.

# --hints--

Пошук-редагування-оновлення елемента має пройти успішно

```js
  const response = await fetch(code + '/_api/find-edit-save', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Poldo',
      age: 40,
      favoriteFoods: ['spaghetti']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Poldo', 'item.name is not what is expected');
  assert.equal(data.age, 40, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['spaghetti', 'hamburger'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 1, 'The item should be previously edited');
```
