---
id: 587d7fb9367417b2b2512c12
title: Об’єднайте помічників пошукових запитів, щоб звузити результати
challengeType: 2
forumTopicId: 301533
dashedName: chain-search-query-helpers-to-narrow-search-results
---

# --description--

Якщо ви не передаєте зворотний виклик як останній аргумент до `Model.find()` (або інших пошукових методів), запит не буде виконуватися. Ви можете зберегти запит в змінній для подальшого використання. Цей тип об’єкту дозволяє створювати запит за допомогою ланцюжкового синтаксису. Сам пошук виконується, коли ви остаточно з’єднали метод `.exec()`. Ви завжди повинні передавати зворотний виклик цьому останньому методу. Існує багато помічників для запитів, тут ми використовуємо найпоширеніші.

# --instructions--

Змініть функцію `queryChain`, щоб знайти людей, які люблять їжу, вказану змінною `foodToSearch`. Відсортуйте їх за `name`, знизьте кількість результатів до двох документів та приховайте їхній вік. Об’єднайте `.find()`, `.sort()`, `.limit()`, `.select()` та `.exec()`. Передайте зворотний виклик `done(err, data)` до `exec()`.

# --hints--

Використання ланцюгових помічників для запитів має бути успішним

```js
  const response = await fetch(code + '/_api/query-tools', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'Pablo', age: 26, favoriteFoods: ['burrito', 'hot-dog'] },
      { name: 'Bob', age: 23, favoriteFoods: ['pizza', 'nachos'] },
      { name: 'Ashley', age: 32, favoriteFoods: ['steak', 'burrito'] },
      { name: 'Mario', age: 51, favoriteFoods: ['burrito', 'prosciutto'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an Array');
  assert.equal(
    data.length,
    2,
    'the data array length is not what expected'
  );
  assert.notProperty(
    data[0],
    'age',
    'The returned first item has too many properties'
  );
  assert.equal(
    data[0].name,
    'Ashley',
    'The returned first item name is not what expected'
  );
  assert.notProperty(
    data[1],
    'age',
    'The returned second item has too many properties'
  );
  assert.equal(
    data[1].name,
    'Mario',
    'The returned second item name is not what expected'
  );
```

