---
id: 587d7fb6367417b2b2512c07
title: створити модель
challengeType: 2
forumTopicId: 301535
dashedName: create-a-model
---

# --description--

**C**RUD Частина I - СТВОРЕННЯ

Перш за все, нам потрібна схема. Кожна схема відповідає колекції MongoDB. Вона визначає форму документів у цій колекції. Схеми є будівельними блоками для моделей. Їх можна вкладати, щоб створювати складні моделі, але в цьому випадку ми залишимо все простим. Модель дозволяє створювати екземпляри ваших об’єктів, які називаються документами.

На серверах взаємодія з базою даних відбувається у функціях-обробниках. Ці функції виконуються, коли відбувається певна подія (наприклад, хтось звертається до кінцевої точки вашого API). Ми дотримуватимемося такого самого підходу в цих вправах. Функція `done()` є функцією зворотного виклику, яка повідомляє, що ми можемо продовжувати після завершення асинхронної операції, такої як вставка, пошук, оновлення або видалення. Вона слідує конвенції Node і має викликатися як `done(null, data)` у разі успіху або `done(err)` у разі помилки.

Увага — при взаємодії з віддаленими сервісами можуть виникати помилки!

```js
/* Example */

const someFunc = function(done) {
  //... do something (risky) ...
  if (error) return done(error);
  done(null, result);
};
```

# --instructions--

Створіть схему для особи з назвою `personSchema` з такою формою:

* Обов’язкове поле `name` типу `String`
* Поле `age` типу `Number`
* Поле `favoriteFoods` типу `[String]`

Використовуйте базові типи схем Mongoose. За бажанням ви можете додати більше полів, використати прості валідатори, такі як required або unique, і встановити значення за замовчуванням. Дивіться нашу <a href="https://www.freecodecamp.org/news/introduction-to-mongoose-for-mongodb-d2a7aa593c57/" target="_blank" rel="noopener noreferrer nofollow">статтю про Mongoose</a>.

Тепер створіть модель на основі `personSchema` і призначте її існуючій змінній `Person`.

# --hints--

Створення екземпляра зі схеми mongoose має бути успішним

```js
  const response = await fetch(code + '/_api/mongoose-model', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Mike',
      age: 28,
      favoriteFoods: ['pizza', 'cheese']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Mike', '"model.name" is not what expected');
  assert.equal(data.age, '28', '"model.age" is not what expected');
  assert.isArray(
    data.favoriteFoods,
    '"model.favoriteFoods" is not an Array'
  );
  assert.include(
    data.favoriteFoods,
    'pizza',
    '"model.favoriteFoods" does not include the expected items'
  );
  assert.include(
    data.favoriteFoods,
    'cheese',
    '"model.favoriteFoods" does not include the expected items'
  );
```
