---
id: 587d7fb6367417b2b2512c06
title: Встановлення та налаштування mongoose
challengeType: 2
forumTopicId: 301540
dashedName: install-and-set-up-mongoose
---

# --description--

Робота над цими завданнями передбачає написання коду одним із таких способів:

- Клонувати <a href="https://github.com/freeCodeCamp/boilerplate-mongomongoose/" target="_blank" rel="noopener noreferrer nofollow">цей репозиторій на GitHub</a> і виконати завдання локально.
- Використати будь-який конструктор сайтів на ваш вибір для виконання проєкту. Обов’язково включіть усі файли з нашого репозиторію на GitHub.

У цьому завданні ви налаштуєте базу даних MongoDB Atlas і імпортуєте потрібні пакети для підключення до неї.

Дотримуйтесь <a href='https://www.freecodecamp.org/news/get-started-with-mongodb-atlas/' target="_blank" rel="noopener noreferrer nofollow">цього посібника</a>, щоб налаштувати хостовану базу даних на MongoDB Atlas.

# --instructions--

`mongoose@^5.11.15` було додано до `package.json` файлу вашого проєкту. Спочатку імпортуйте mongoose як `mongoose` у `myApp.js`. Далі створіть `.env` файл і додайте в нього змінну `MONGO_URI`. Її значенням має бути URI вашої бази даних MongoDB Atlas. Обов’язково оточіть URI одинарними або подвійними лапками, і пам’ятайте, що у змінних середовища не можна використовувати пробіли навколо `=`. Наприклад, `MONGO_URI='VALUE'`.

Коли закінчите, підключіться до бази даних, викликавши метод `connect` у вашому `myApp.js` файлі за допомогою такого синтаксису:

```js
mongoose.connect(<Your URI>, { useNewUrlParser: true, useUnifiedTopology: true });
```

# --hints--

У файлі package.json має бути залежність "mongoose version ^5.11.15"

```js
  const response = await fetch(code + '/_api/file/package.json');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  var packJson = JSON.parse(data);
  assert.property(packJson.dependencies, 'mongoose');
  assert.match(
    packJson.dependencies.mongoose,
    /^\^5\.11\.15/,
    'Wrong version of "mongoose". It should be ^5.11.15'
  );
```

"mongoose" має бути підключений до бази даних

```js
  const response = await fetch(code + '/_api/is-mongoose-ok');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(data.isMongooseOk, 'mongoose is not connected');
```
