---
id: 587d7fb9367417b2b2512c12
title: Зв’язування допоміжних запитів для звуження результатів пошуку
challengeType: 2
forumTopicId: 301533
dashedName: chain-search-query-helpers-to-narrow-search-results
---

# --description--

Якщо ви не передасте функцію зворотного виклику як останній аргумент у `Model.find()` (або в інші методи пошуку), запит не буде виконано. Ви можете зберегти запит у змінній для подальшого використання. Такий об’єкт дозволяє будувати запит, використовуючи синтаксис зв’язування. Фактичний пошук у базі даних виконується, коли ви нарешті зв’язуєте метод `.exec()`. Ви завжди повинні передавати свою функцію зворотного виклику цьому останньому методу. Існує багато допоміжних запитів, тут ми використаємо найпоширеніші.

# --instructions--

Змініть функцію `queryChain`, щоб знайти людей, які люблять їжу, вказану у змінній `foodToSearch`. Відсортуйте їх за `name`, обмежте результати двома документами та приховайте їхній вік. Зв’яжіть `.find()`, `.sort()`, `.limit()`, `.select()`, а потім `.exec()`. Передайте функцію зворотного виклику `done(err, data)` у `exec()`.

# --hints--

Зв’язування допоміжних запитів має працювати

```js
  const response = await fetch(code + '/_api/query-tools', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'Pablo', age: 26, favoriteFoods: ['burrito', 'hot-dog'] },
      { name: 'Bob', age: 23, favoriteFoods: ['pizza', 'nachos'] },
      { name: 'Ashley', age: 32, favoriteFoods: ['steak', 'burrito'] },
      { name: 'Mario', age: 51, favoriteFoods: ['burrito', 'prosciutto'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an Array');
  assert.equal(
    data.length,
    2,
    'the data array length is not what expected'
  );
  assert.notProperty(
    data[0],
    'age',
    'The returned first item has too many properties'
  );
  assert.equal(
    data[0].name,
    'Ashley',
    'The returned first item name is not what expected'
  );
  assert.notProperty(
    data[1],
    'age',
    'The returned second item has too many properties'
  );
  assert.equal(
    data[1].name,
    'Mario',
    'The returned second item name is not what expected'
  );
```
