---
id: 587d7fb7367417b2b2512c0d
title: Використання model.findById() для пошуку у базі даних за _id
challengeType: 2
forumTopicId: 301544
dashedName: use-model-findbyid-to-search-your-database-by-id
---

# --description--

Під час збереження документа MongoDB автоматично додає поле `_id` і встановлює його унікальним алфавітно-цифровим ключем. Пошук за `_id` — це надзвичайно часта операція, тому Mongoose надає спеціальний метод для цього.

# --instructions--

Змініть `findPersonById`, щоб знайти єдину людину з заданим `_id`, використовуючи `Model.findById() -> Person`. Використайте аргумент функції `personId` як ключ пошуку.

# --hints--

Пошук елемента за Id має бути успішним

```js
  const response = await fetch(code + '/_api/find-by-id');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'test', 'item.name is not what expected');
  assert.equal(data.age, 0, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['none'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```
