---
id: 587d7fb7367417b2b2512c0c
title: Використання model.findOne() для повернення одного відповідного документа з вашої бази даних
challengeType: 2
forumTopicId: 301545
dashedName: use-model-findone-to-return-a-single-matching-document-from-your-database
---

# --description--

`Model.findOne()` працює так само, як `Model.find()`, але повертає лише один документ (не масив), навіть якщо є кілька елементів. Це особливо корисно при пошуку за пропами, які ви оголосили унікальними.

# --instructions--

Змініть функцію `findOneByFood`, щоб знайти лише одну людину, у якої в улюблених стравах є певна їжа, використовуючи `Model.findOne() -> Person`. Використовуйте аргумент функції `food` як ключ пошуку.

# --hints--

Пошук одного елемента має бути успішним

```js
  const response = await fetch(code + '/_api/find-one-by-food', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Gary',
      age: 46,
      favoriteFoods: ['chicken salad']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Gary', 'item.name is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['chicken salad'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```
