---
id: 587d7fb8367417b2b2512c11
title: Видалити багато документів за допомогою model.remove()
challengeType: 2
forumTopicId: 301538
dashedName: delete-many-documents-with-model-remove
---

# --description--

`Model.remove()` корисна для видалення всіх документів, що відповідають заданим критеріям.

# --instructions--

Змініть функцію `removeManyPeople`, щоб видалити всіх людей, чиє ім’я міститься у змінній `nameToRemove`, використовуючи `Model.remove()`. Передайте це у документ запиту з полем `name`, а також функцію зворотного виклику.

**Примітка:** `Model.remove()` не повертає видалений документ, а JSON-об’єкт, що містить результат операції та кількість затронутих елементів. Не забудьте передати його у функцію зворотного виклику `done()`, оскільки ми використовуємо це в тестах.

# --hints--

Видалення багатьох елементів одночасно має пройти успішно

```js
  const response = await fetch(code + '/_api/remove-many-people', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'Mary', age: 16, favoriteFoods: ['lollipop'] },
      { name: 'Mary', age: 21, favoriteFoods: ['steak'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(!!data.ok, 'The mongo stats are not what expected');
  assert.equal(
    data.n,
    2,
    'The number of items affected is not what expected'
  );
  assert.equal(data.count, 0, 'the db items count is not what expected');
```
