---
id: 587d7fb7367417b2b2512c0a
title: Створення багатьох записів за допомогою model.create()
challengeType: 2
forumTopicId: 301537
dashedName: create-many-records-with-model-create
---

# --description--

Іноді потрібно створити багато екземплярів ваших моделей, наприклад, при заповненні бази даних ініціальними даними. `Model.create()` приймає масив об’єктів, як `[{name: 'John', ...}, {...}, ...]`, як перший аргумент і зберігає їх усіх у базі даних.

# --instructions--

Змініть функцію `createManyPeople`, щоб створити багато людей за допомогою `Model.create()` з аргументом `arrayOfPeople`.

**Примітка:** Ви можете повторно використати модель, яку створили в попередньому завданні.

# --hints--

Створення багатьох елементів бази даних одночасно має бути успішним

```js
  const response = await fetch(code + '/_api/create-many-people', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'John', age: 24, favoriteFoods: ['pizza', 'salad'] },
      { name: 'Mary', age: 21, favoriteFoods: ['onions', 'chicken'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an array');
  assert.equal(
    data.length,
    2,
    'the response does not contain the expected number of items'
  );
  assert.equal(data[0].name, 'John', 'The first item is not correct');
  assert.equal(
    data[0].__v,
    0,
    'The first item should be not previously edited'
  );
  assert.equal(data[1].name, 'Mary', 'The second item is not correct');
  assert.equal(
    data[1].__v,
    0,
    'The second item should be not previously edited'
  );
```
