---
id: 587d7b85367417b2b2512b39
title: Fehlende öffnende und schließende runde Klammern nach einem Funktionsaufruf erkennen
challengeType: 1
forumTopicId: 301185
dashedName: catch-missing-open-and-closing-parenthesis-after-a-function-call
---

# --description--

Wenn eine Funktion oder Methode keine Argumente erwartet, kann es passieren, dass Sie vergessen, die (leeren) öffnenden und schließenden runden Klammern beim Aufrufen einzufügen. Häufig wird das Ergebnis eines Funktionsaufrufs in einer Variablen gespeichert, um es später im Code weiterzuverwenden. Dieser Fehler lässt sich erkennen, indem Sie die Werte der Variablen (oder deren Typen) in der Konsole ausgeben und feststellen, dass eine Variable auf eine Funktionsreferenz gesetzt ist, anstatt auf den erwarteten Wert, den die Funktion zurückgibt.

Die Variablen im folgenden Beispiel sind unterschiedlich:

```js
function myFunction() {
  return "You rock!";
}
let varOne = myFunction;
let varTwo = myFunction();
```

Hier ist `varOne` die Funktion `myFunction`, und `varTwo` ist der String `You rock!`.

# --instructions--

Beheben Sie den Code so, dass die Variable `result` auf den Wert gesetzt wird, der durch den Aufruf der Funktion `getNine` zurückgegeben wird.

# --hints--

Ihr Code sollte die Variable `result` so korrigieren, dass sie auf die Zahl gesetzt wird, die die Funktion `getNine` zurückgibt.

```js
assert(result == 9);
```

Ihr Code sollte die Funktion `getNine` aufrufen.

```js
assert(__helpers.removeJSComments(code).match(/getNine\(\)/g).length == 2);
```

# --seed--

## --seed-contents--

```js
function getNine() {
  let x = 6;
  let y = 3;
  return x + y;
}

let result = getNine;
console.log(result);
```

# --solutions--

```js
function getNine() {
 let x = 6;
 let y = 3;
 return x + y;
}

let result = getNine();
console.log(result);
```
