---
id: 587d7b86367417b2b2512b3b
title: Off-by-one-Fehler beim Verwenden von Indexierung erkennen
challengeType: 1
forumTopicId: 301189
dashedName: catch-off-by-one-errors-when-using-indexing
---

# --description--

<dfn>Off-by-one-Fehler</dfn> (manchmal OBOE genannt) treten auf, wenn Sie versuchen, einen bestimmten Index eines Strings oder Arrays anzusprechen (um einen Abschnitt auszuschneiden oder zuzugreifen) oder wenn Sie über die Indizes davon iterieren. Die Indexierung in JavaScript beginnt bei null, nicht bei eins, was bedeutet, dass der letzte Index immer um eins kleiner als die Länge des Elements ist. Wenn Sie versuchen, auf einen Index zuzugreifen, der der Länge entspricht, kann das Programm einen „Index außerhalb des Bereichs“-Referenzfehler werfen oder `undefined` ausgeben.

Wenn Sie String- oder Array-Methoden verwenden, die Indexbereiche als Argumente nehmen, hilft es, die Dokumentation zu lesen und zu verstehen, ob diese inklusive sind (das Element am angegebenen Index ist Teil des Rückgabewerts) oder nicht. Hier sind einige Beispiele für Off-by-one-Fehler:

```js
let alphabet = "abcdefghijklmnopqrstuvwxyz";
let len = alphabet.length;
for (let i = 0; i <= len; i++) {
  console.log(alphabet[i]);
}
for (let j = 1; j < len; j++) {
  console.log(alphabet[j]);
}
for (let k = 0; k < len; k++) {
  console.log(alphabet[k]);
}
```

Das erste Beispiel hier durchläuft die Schleife eine zu viele Male, und das zweite durchläuft sie eine zu wenige Male (der erste Index, 0, wird ausgelassen). Das dritte Beispiel ist korrekt.

# --instructions--

Beheben Sie die beiden Indexierungsfehler in der folgenden Funktion, sodass alle Zahlen von 1 bis 5 in der Konsole ausgegeben werden.

# --hints--

Ihr Code sollte die Anfangsbedingung der Schleife so setzen, dass sie beim ersten Index startet.

```js
assert(__helpers.removeJSComments(code).match(/i\s*?=\s*?0\s*?;/g).length == 1);
```

Ihr Code sollte die Anfangsbedingung der Schleife so korrigieren, dass der Index bei 0 beginnt.

```js
assert(!__helpers.removeJSComments(code).match(/i\s?=\s*?1\s*?;/g));
```

Ihr Code sollte die Endbedingung der Schleife so setzen, dass sie beim letzten Index stoppt.

```js
assert(__helpers.removeJSComments(code).match(/i\s*<\s*len\s*;|i\s*<=\s*len\s*-\s*1\s*;/g).length == 1);
```

Ihr Code sollte die Endbedingung der Schleife so korrigieren, dass sie einen Index vor der Länge stoppt.

```js
assert(!__helpers.removeJSComments(code).match(/i\s*?<=\s*?len;/g));
```

# --seed--

## --seed-contents--

```js
function countToFive() {
  let firstFive = "12345";
  let len = firstFive.length;
  // Only change code below this line
  for (let i = 1; i <= len; i++) {
  // Only change code above this line
    console.log(firstFive[i]);
  }
}

countToFive();
```

# --solutions--

```js
function countToFive() {
 let firstFive = "12345";
 let len = firstFive.length;
 // Only change code below this line
 for (let i = 0; i < len; i++) {
 // Only change code above this line
   console.log(firstFive[i]);
 }
}

countToFive();
```
