---
id: 587d7b85367417b2b2512b3a
title: "Argumente abfangen, die beim Aufrufen einer Funktion in falscher Reihenfolge übergeben werden"
challengeType: 1
forumTopicId: 301184
dashedName: catch-arguments-passed-in-the-wrong-order-when-calling-a-function
---

# --description--

Im Anschluss an die Diskussion über das Aufrufen von Funktionen ist der nächste Bug, auf den Sie achten sollten, wenn die Argumente einer Funktion in der falschen Reihenfolge übergeben werden. Wenn die Argumente unterschiedliche Typen haben, zum Beispiel wenn eine Funktion ein Array und eine ganze Zahl erwartet, führt dies wahrscheinlich zu einem Laufzeitfehler. Sind die Argumente vom gleichen Typ (beispielsweise alle ganze Zahlen), ergibt die Logik des Codes keinen Sinn. Stellen Sie sicher, dass Sie alle erforderlichen Argumente in der richtigen Reihenfolge übergeben, um diese Probleme zu vermeiden.

# --instructions--

Die Funktion `raiseToPower` potenziert eine Basis mit einem Exponenten. Leider wird sie nicht korrekt aufgerufen – beheben Sie den Code so, dass der Wert von `power` wie erwartet 8 ist.

# --hints--

Ihr Code sollte die Variable `power` so korrigieren, dass sie 2 hoch 3 entspricht und nicht 3 hoch 2.

```js
assert(power == 8);
```

Ihr Code sollte die richtige Reihenfolge der Argumente beim Aufruf der Funktion `raiseToPower` verwenden.

```js
assert(__helpers.removeJSComments(code).match(/raiseToPower\(\s*?base\s*?,\s*?exp\s*?\);/g));
```

# --seed--

## --seed-contents--

```js
function raiseToPower(b, e) {
  return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(exp, base);
console.log(power);
```

# --solutions--

```js
function raiseToPower(b, e) {
 return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(base, exp);
console.log(power);
```
