---
id: 587d7b86367417b2b2512b3c
title: Vorsicht beim erneuten Initialisieren von Variablen innerhalb einer Schleife
challengeType: 1
forumTopicId: 301194
dashedName: use-caution-when-reinitializing-variables-inside-a-loop
---

# --description--

Manchmal ist es notwendig, Informationen zu speichern, Zähler zu erhöhen oder Variablen innerhalb einer Schleife zurückzusetzen. Ein mögliches Problem entsteht, wenn Variablen entweder hätten neu initialisiert werden müssen und es nicht wurden oder umgekehrt. Das ist besonders gefährlich, wenn Sie versehentlich die Variable zurücksetzen, die für die Abbruchbedingung verwendet wird, was zu einer Endlosschleife führt.

Das Ausgeben von Variablenwerten bei jedem Durchlauf Ihrer Schleife mit `console.log()` kann fehlerhaftes Verhalten aufdecken, das mit dem Zurücksetzen oder Nicht-Zurücksetzen einer Variable zusammenhängt.

# --instructions--

Die folgende Funktion soll ein zweidimensionales Array mit `m` Zeilen und `n` Spalten aus Nullen erzeugen. Leider liefert sie nicht das erwartete Ergebnis, weil die `row`-Variable in der äußeren Schleife nicht neu initialisiert (wieder auf ein leeres Array gesetzt) wird. Korrigieren Sie den Code so, dass er ein korrektes 3x2-Array aus Nullen zurückgibt, das wie `[[0, 0], [0, 0], [0, 0]]` aussieht.

# --hints--

Ihr Code sollte die `matrix`-Variable auf ein Array setzen, das 3 Zeilen mit jeweils 2 Spalten aus Nullen enthält.

```js
assert(JSON.stringify(matrix) == '[[0,0],[0,0],[0,0]]');
```

Die `matrix`-Variable sollte 3 Zeilen haben.

```js
assert(matrix.length == 3);
```

Die `matrix`-Variable sollte in jeder Zeile 2 Spalten haben.

```js
assert(
  matrix[0].length == 2 && matrix[1].length === 2 && matrix[2].length === 2
);
```

`zeroArray(4,3)` sollte ein Array mit 4 Zeilen und 3 Spalten aus Nullen zurückgeben.

```js
assert(JSON.stringify(zeroArray(4,3)) == '[[0,0,0],[0,0,0],[0,0,0],[0,0,0]]');
```

# --seed--

## --seed-contents--

```js
function zeroArray(m, n) {
  // Creates a 2-D array with m rows and n columns of zeroes
  let newArray = [];
  let row = [];
  for (let i = 0; i < m; i++) {
    // Adds the m-th row into newArray

    for (let j = 0; j < n; j++) {
      // Pushes n zeroes into the current row to create the columns
      row.push(0);
    }
    // Pushes the current row, which now has n zeroes in it, to the array
    newArray.push(row);
  }
  return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```

# --solutions--

```js
function zeroArray(m, n) {
 // Creates a 2-D array with m rows and n columns of zeroes
 let newArray = [];
 for (let i = 0; i < m; i++) {
   let row = [];
   // Adds the m-th row into newArray

   for (let j = 0; j < n; j++) {
     // Pushes n zeroes into the current row to create the columns
     row.push(0);
   }
   // Pushes the current row, which now has n zeroes in it, to the array
   newArray.push(row);
 }
 return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```
