---
id: 587d7b84367417b2b2512b36
title: 'Ungeöffnete runde Klammern, Klammern, geschweifte Klammern und Anführungszeichen erkennen'
challengeType: 1
forumTopicId: 301190
dashedName: catch-unclosed-parentheses-brackets-braces-and-quotes
---

# --description--

Ein weiterer Syntaxfehler, auf den Sie achten sollten, ist, dass alle öffnenden runden Klammern, Klammern, geschweiften Klammern und Anführungszeichen ein schließendes Gegenstück haben. Das Vergessen eines Teils passiert häufig, wenn Sie bestehenden Code bearbeiten und Elemente mit einem der Paar-Typen einfügen. Achten Sie auch darauf, wenn Sie Codeblöcke in andere verschachteln, zum Beispiel wenn Sie eine Callback-Funktion als Argument an eine Methode übergeben.

Eine Möglichkeit, diesen Fehler zu vermeiden, besteht darin, sobald das öffnende Zeichen eingegeben wird, sofort das passende schließende Zeichen einzufügen, dann den Cursor zurück zwischen die beiden zu setzen und mit dem Codieren fortzufahren. Glücklicherweise erzeugen die meisten modernen Code-Editoren die zweite Hälfte des Paares automatisch.

# --instructions--

Beheben Sie die zwei Fehler bei den Paaren im Code.

# --hints--

Ihr Code sollte das fehlende Teil des Arrays beheben.

```js
assert(__helpers.removeJSComments(code).match(/myArray\s*?=\s*?\[\s*?1\s*?,\s*?2\s*?,\s*?3\s*?\];/g));
```

Ihr Code sollte das fehlende Teil der `.reduce()`-Methode beheben. Die Ausgabe in der Konsole sollte zeigen, dass `Sum of array values is: 6`.

```js
assert(arraySum === 6);
```

# --seed--

## --seed-contents--

```js
let myArray = [1, 2, 3;
let arraySum = myArray.reduce((previous, current =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```

# --solutions--

```js
let myArray = [1, 2, 3];
let arraySum = myArray.reduce((previous, current) =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```
