---
id: 66ec4c8e9878d8441956516f
title: 创建一个图书目录表格
challengeType: 25
dashedName: build-a-book-catalog-table
demoType: onClick
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该创建一个用于列出图书信息的 `table` 元素。
1. 你的表中应该有一个一行的表头元素。
1. 你的表头元素中的行应该有四个表标题元素，文本依次是 `Title`、`Author`、`Genre` 和 `Publication Year`。
1. 你的表中应该有一个至少是五行的表体元素。
1. 你的表体中的每一行都应该有四个表数据元素，显示图书的名称、作者、体裁和出版年份。
1. 你的表格应该有一个一行的表脚元素。
1. 你的表脚元素中的行应当包含一个表格数据单元格元素，该元素需横跨四列，并且显示文本 `Total Books: N` ，其中 `N` 应替换为你的表格中的图书数量。

# --hints--

你应该有一个 `table` 元素。

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

在你的 `table` 元素中，应该有一个 `thead` 元素。

```js
assert.lengthOf(document.querySelectorAll('table thead'), 1);
```

在你的 `thead` 元素内，应该有一个 `tr` 元素。

```js
assert.lengthOf(document.querySelectorAll('thead tr'), 1);
```

在你的 `thead` 元素的行中，应该有四个 `th` 元素。

```js
assert.lengthOf(document.querySelectorAll('thead tr th'), 4);
```

你这四个 `th` 元素应依次显示文本 `Title`、`Author`、`Genre` 和 `Publication Year` 。

```js
const ths = document.querySelectorAll('thead tr th');
assert.equal(ths[0]?.textContent, 'Title');
assert.equal(ths[1]?.textContent, 'Author');
assert.equal(ths[2]?.textContent, 'Genre');
assert.equal(ths[3]?.textContent, 'Publication Year');
```

在你的 `table` 元素中，应该有一个 `tbody` 元素。

```js
assert.lengthOf(document.querySelectorAll('table tbody'), 1);
```

你的 `tbody` 元素至少应包含五行内容。

```js
assert.isAtLeast(document.querySelectorAll('tbody tr').length, 5);
```

你的 `tbody` 元素中的每一行都应恰好有四个 `td` 元素作为其子元素。

```js
const rows = document.querySelectorAll('tbody tr');
assert.isAtLeast(rows.length, 1);

rows.forEach(row => {
  assert.lengthOf(row.querySelectorAll('td'), 4);
  assert.lengthOf(row.children, 4);
});
```

你的表体中的每个 `td` 元素都应包含有关图书的信息。

```js
const tds = document.querySelectorAll('tbody tr td');
assert.isAtLeast(tds.length, 1);

tds.forEach(td => {
  assert.isAtLeast(td.textContent.length, 1);
});
```

在你的 `table` 元素内，应当有一个 `tfoot` 元素。

```js
assert.lengthOf(document.querySelectorAll('table tfoot'), 1);
```

在你的 `tfoot` 元素中，应该恰好只有一个 `tr` 元素。

```js
assert.lengthOf(document.querySelectorAll('tfoot tr'), 1);
```

你的 `tfoot` 元素行中的 `td` 元素应将其 `colspan` 属性设置为 `4` 。

```js
assert.equal(document.querySelector('tfoot tr td')?.colSpan, 4);
```

你的 `tfoot` 元素行中的 `td` 元素应当显示文本 `Total Books: [number of books in your table]` 。

```js
const numberOfBooks = document.querySelectorAll('tbody tr')?.length;
assert.equal(document.querySelector('tfoot tr td').textContent, `Total Books: ${numberOfBooks}`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Book Catalog</title>
</head>

<body>
</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Book Catalog</title>
</head>
<body>
    <h1>Book Catalog</h1>
  <table>
    <thead>
      <tr>
        <th>Title</th>
        <th>Author</th>
        <th>Genre</th>
        <th>Publication Year</th>
      </tr>
    </thead>
    <tbody>
      <tr>
        <td>How to Contribute to Open-Source Projects – A Handbook for Beginners</td>
        <td>Hillary Nyakundi</td>
        <td>Open Source</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>Learn Linux for Beginners: From Basics to Advanced Techniques</td>
        <td>Zaira Hira</td>
        <td>Linux</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>How to Learn to Code and Get a Developer Job</td>
        <td>Quincy Larson</td>
        <td>Learn To Code</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>The Regular Expressions Book – RegEx for JavaScript Developers</td>
        <td>Kolade Chris</td>
        <td>Regular Expressions</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>The Python Code Example Handbook</td>
        <td>Farhan Hasin Chowdhury</td>
        <td>Python</td>
        <td>2023</td>
      </tr>
    </tbody>
    <tfoot>
      <tr>
        <td colspan="4">Total Books: 5</td>
      </tr>
    </tfoot>
  </table>
</body>
</html>
```
