---
id: 615f2d4150fe0d4cbd0f2628
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Під вашим елементом `h1` додайте елемент `p` з текстом `8 servings per container`.

# --hints--

Ви повинні додати новий елемент `p`.

```js
assert.exists(document.querySelector('p'));
```

Ваш елемент `p` має бути всередині вашого елемента `body`.

```js
assert.equal(document.querySelector('p')?.parentElement?.localName, 'body');
```

Ваш елемент `p` має йти після вашого елемента `h1`.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.localName, 'h1');
```

Ваш елемент `p` має містити текст `8 servings per container`.

```js
assert.equal(document.querySelector('p')?.innerText, '8 servings per container');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
    
</body>
</html>
--fcc-editable-region--
```

```css

```
