---
id: 615f41c979787462e76dab90
title: Крок 20
challengeType: 0
dashedName: step-20
---

# --description--

Ваш новий клас поки що не має стилів. Створіть селектор `.bold` і надайте йому множину властивостей `font-weight` зі значенням `800`, щоб зробити текст жирним.

Також видаліть властивість `font-weight` з вашого селектора `h1`.

# --hints--

У вас має бути селектор `.bold`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.bold'));
```

Ваш селектор `.bold` має мати множину властивостей `font-weight` зі значенням `800`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.bold')?.fontWeight === '800');
```

Ваш селектор `h1` не повинен мати властивість `font-weight`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === "");
```

Ви не повинні видаляти свій селектор `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <div class="divider"></div>
    <p>8 servings per container</p>
    <p class="bold">Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--
h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px;
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

--fcc-editable-region--
```
