---
id: 615f34ecc1091b4fd5a8a484
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Усередині вашого елемента `head` додайте елемент `link` з атрибутом `rel`, встановленим у `stylesheet`, та атрибутом `href`, встановленим у `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

Це імпортує сімейство шрифтів `Open Sans` з ваговими значеннями шрифту `400`, `700` і `800`.

Також додайте елемент `link`, щоб зв’язати ваш файл `styles.css`.

# --hints--

У вашому коді має бути два елементи `link`.

```js
assert.strictEqual(document.querySelectorAll('link').length, 2);
```

Ваші елементи `link` мають бути порожніми елементами, вони не повинні мати закриваючого тегу `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

Ваші два елементи `link` мають бути всередині елемента `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link').length, 2);
```

Обидва ваші елементи `link` мають мати атрибут `rel`, встановлений у `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Один із ваших елементів `link` має мати атрибут `href`, встановлений у `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
```

Один із ваших елементів `link` має мати атрибут `href`, встановлений у `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
