---
id: 61329d52e5010e08d9b9d66b
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Ein weiteres wichtiges `meta`-Element für Barrierefreiheit und SEO ist die `description`-Definition. Der Wert des `content`-Attributs wird von Suchmaschinen verwendet, um eine Beschreibung Ihrer Seite bereitzustellen.

Fügen Sie ein `meta`-Element mit dem `name`-Attribut, das auf `description` gesetzt ist, hinzu, und geben Sie ihm ein nützliches `content`-Attribut.

# --hints--

Sie sollten dem `meta` ein neues `head`-Element hinzufügen.

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

Sie sollten dem `meta` ein `name`-Attribut mit dem Wert `description` geben.

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

Sie sollten dem `meta` ein `content`-Attribut geben.

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

Der Wert des `content`-Attributs sollte nicht mehr als 165 Zeichen umfassen. _Dies ist die maximale Beschreibungs­länge von Google._

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
