---
id: 61329b210dac0b08047fd6ab
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

In Fortsetzung mit den `meta`-Elementen gibt eine `viewport`-Definition dem Browser vor, wie die Seite gerendert wird. Das Einfügen einer solchen verbessert die visuelle Barrierefreiheit auf Mobilgeräten und steigert die _SEO_ (Suchmaschinenoptimierung).

Fügen Sie eine `viewport`-Definition mit einem `content`-Attribut hinzu, das die `width` und `initial-scale` der Seite beschreibt.

# --hints--

Sie sollten ein weiteres `meta`-Element im `head` erstellen.

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

Sie sollten dem `meta` ein `name`-Attribut mit dem Wert `viewport` geben.

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

Sie sollten dem `meta` ein `content`-Attribut mit dem Wert `width=device-width, initial-scale=1` geben.

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
