---
id: 614ccc21ea91ef1736b9b578
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Willkommen zum ersten Teil des Barrierefreiheits-Quiz. Da Sie sich zu einer erfahrenen HTML- und CSS-Entwicklerin bzw. einem erfahrenen Entwickler entwickeln, haben wir mit dem grundlegenden Boilerplate-Code für Sie begonnen.

Starten Sie diese Barrierefreiheitsreise, indem Sie Ihrem `lang`-Element ein `html`-Attribut hinzufügen. Dies hilft Screenreadern dabei, die Sprache der Seite zu erkennen.

# --hints--

Sie sollten dem `html`-Element ein `lang`-Attribut geben. _Hinweis: Für Englisch können Sie den Wert `en` verwenden._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
