---
id: 68c3144fc395d681180cde15
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

页面上使用标题似乎仍存在一些问题。如果你查看第一个 `div` 元素，会发现有两个 `h4` 元素用作子标题。更好的做法是为这些子标题使用 `h3` 元素，因为它们比 `h2` 元素低一级。

将两个 `h4` 元素更改为 `h3` 元素。

# --hints--

你的 `h4` 元素中带有 `Early Challenges` 文本的部分应为 `h3` 元素。

```js
const post1Subheading = document.querySelector("#post1 + p + h3");
assert.equal(post1Subheading?.tagName, "H3");
assert.equal(post1Subheading?.textContent, "Early Challenges");
```

你的 `h4` 元素中带有 `Breakthroughs` 文本的部分应为 `h3` 元素。

```js
const post1SecondSubheading = document.querySelector("#post1 + p + h3 + p + h3");
assert.equal(post1SecondSubheading?.tagName, "H3");
assert.equal(post1SecondSubheading?.textContent, "Breakthroughs");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Camperbot's Blog</title>
</head>
<body>
  <h1>Welcome to Camperbot's Blog</h1>

  <h2>Navigation</h2>
  <ul>
    <li><a href="#post1">My Journey</a></li>
    <li><a href="#post2">Accessibility</a></li>
    <li><a href="#post3">Next Steps</a></li>
  </ul>
--fcc-editable-region--
  <div>
    <h2 id="post1">My Journey Learning to Code</h2>
    <p>I started learning to code a few months ago and it's been a wild ride!</p>
    
    <h4>Early Challenges</h4>
    <p>At first, syntax was really confusing.</p>
    
    <h4>Breakthroughs</h4>
    <p>Eventually things started to click.</p>
  </div>
--fcc-editable-region--
  <div>
    <h2 id="post2">Accessibility Matters</h2>
    <p>Today I learned that not everyone uses the web the same way I do.</p>

    <h5>Screen Readers</h5>
    <p>These tools help visually impaired users browse websites.</p>
  </div>

  <div>
    <h2 id="post3">What's Next?</h2>
    <p>I'm excited to dive into JavaScript and build interactive features!</p>

    <h3>Coming soon: My first JavaScript project!</h3>
    <p>Stay tuned for some exciting interactive blog features.</p>
  </div>
  <h3>Contact Me</h3>
  <p>Email me at camperbot@blog.io</p>
</body>
</html>
```
