---
id: 68c31dd917c3ae855e8728c7
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

在研讨会的最后部分，底部的联系部分需要进行一些更改。

第一个更改是将联系部分包裹在一个 `footer` 元素内。这将帮助屏幕阅读器将此部分识别为页面的页脚。

# --hints--

你的页面应包含一个 `footer` 元素。

```js
const footerElement = document.querySelector("footer");
assert.isNotNull(footerElement);
```

你的 `footer` 元素应包含一个带有文本 `Contact Me` 的 `h3` 元素。

```js
const h3Element = document.querySelector("footer h3");
assert.isNotNull(h3Element);
assert.strictEqual(h3Element?.textContent, "Contact Me");
```

你的 `footer` 元素应包含一个带有文本 `Email me at camperbot@blog.io` 的 `p` 元素。

```js
const pElement = document.querySelector("footer p");
assert.isNotNull(pElement);
assert.strictEqual(pElement?.textContent, "Email me at camperbot@blog.io");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Camperbot's Blog</title>
</head>
<body>
  <h1>Welcome to Camperbot's Blog</h1>

  <nav>
    <h2>Navigation</h2>
    <ul>
      <li><a href="#post1">My Journey</a></li>
      <li><a href="#post2">Accessibility</a></li>
      <li><a href="#post3">Next Steps</a></li>
    </ul>
  </nav>

  <main>
    <article>
      <h2 id="post1">My Journey Learning to Code</h2>
      <p>I started learning to code a few months ago and it's been a wild ride!</p>
      
      <h3>Early Challenges</h3>
      <p>At first, syntax was really confusing.</p>
      
      <h3>Breakthroughs</h3>
      <p>Eventually things started to click.</p>
    </article>

    <article>
      <h2 id="post2">Accessibility Matters</h2>
      <p>Today I learned that not everyone uses the web the same way I do.</p>

      <h3>Screen Readers</h3>
      <p>These tools help visually impaired users browse websites.</p>
    </article>

    <article>
      <h2 id="post3">What's Next?</h2>
      <p>I'm excited to dive into JavaScript and build interactive features!</p>

      <h3>Coming soon: My first JavaScript project!</h3>
      <p>Stay tuned for some exciting interactive blog features.</p>
    </article>
  </main>
--fcc-editable-region--
  <h3>Contact Me</h3>
  <p>Email me at camperbot@blog.io</p>
--fcc-editable-region--
</body>
</html>
```
