---
id: 6565a536ba1f9f25bd30e88b
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 프로젝트에서는 Luhn 알고리즘을 구현할 것입니다. 이 알고리즘은 다양한 식별 번호를 검증하는 공식입니다.

`main`라는 함수를 선언하는 것부터 시작하세요. 이 함수가 프로그램의 진입점 역할을 합니다. 오류를 피하려면 `pass` 키워드를 사용하세요.

# --hints--

코드에 `def main():`가 있어야 합니다.

```js
({
    test: () => {
        assert(runPython(`
        import inspect
        inspect.isfunction(main)
        `))
    }
})
```

코드에 `pass`가 있어야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const pass = __helpers.python.getDef("\n"+transformedCode, "main");
        assert.include(pass?.function_body, "pass");
    }
})
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
