---
id: 65687946130b0ea10aa19b75
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

스크립트 끝에서 `main` 함수를 호출하세요.

# --hints--

`main()`이 `main` 함수 밖에 있어야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        assert.match(transformedCode, /\nmain\(\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)


--fcc-editable-region--
```
