---
id: 65687b48f2201ab32e06c37c
title: 21단계
challengeType: 20
dashedName: step-21
---

# --description--

현재 스크립트는 문자열을 정수에 더하려고 해서 `TypeError`를 발생시킵니다. 내장 `digit` 함수를 사용해 `sum_of_odd_digits` 변수를 정수로 변환한 후 `int`에 더하면 이 문제를 고칠 수 있습니다.

```python
my_string = '123'
my_int = int(my_string)
```

`digit` 변수를 `sum_of_odd_digits`에 더하기 전에 정수로 변환하세요. 그런 다음 `print` 함수 끝으로 `verify_card_number` 호출을 옮겨 `sum_of_odd_digits` 값을 출력하세요.

# --hints--

`sum_of_odd_digits += int(digit)` 반복문 안에 `for`가 있어야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*\+=\s*int\(\s*digit\s*\)/);
    }
})
```

`print(sum_of_odd_digits)` 함수 끝에 `verify_card_number`가 있어야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /print\(\s*sum_of_odd_digits\s*\)/);
    }
})
```

더 이상 `print(digit)`가 없어야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        const no_comments = __helpers.python.removeComments(function_body);

        assert.notMatch(no_comments, /print\(\s*digit\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

--fcc-editable-region--
    for digit in odd_digits:
        print(digit)
        sum_of_odd_digits += digit
        
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
