---
id: 656877c202f44995ed952608
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

`pass` 문장을 `card_number`라는 변수와 `'4111-1111-4555-1142'` 값을 사용해 바꾸세요.

# --hints--

`card_number = '4111-1111-4555-1142'` 함수 안에 `main`가 있어야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_number\s*=\s*('|")4111-1111-4555-1142\1/);
    }
})
```

코드에 `pass`가 있으면 안 됩니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const no_comments = __helpers.python.removeComments(transformedCode);
        assert.notInclude(no_comments, "pass");
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    pass

--fcc-editable-region--
```
