---
id: 65687b182c2a8fb1470e0b97
title: 20단계
challengeType: 20
dashedName: step-20
---

# --description--

`for` 반복문 안에서 `+=` 연산자를 사용해 `digit`를 `sum_of_odd_digits` 변수에 더하세요.

이렇게 하면 문자열을 정수에 더하려고 해서 스크립트가 `TypeError`를 발생시키지만 걱정하지 마세요. 다음 단계에서 이를 해결하는 방법을 더 배우게 됩니다.

# --hints--

`sum_of_odd_digits += digit` 반복문 안에 `for`가 있어야 합니다.

```js
const transformedCode = code.replace(/\r/g, "");
const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
const { function_body } = verify_card_number;

assert.match(function_body, /sum_of_odd_digits\s*\+=\s*digit/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        print(digit)
        
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
