---
id: 656879a66338b2a461d5d307
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

<!-- TODO: Find better places to split explanation up. -->

Luhn 알고리즘은 다음과 같습니다:

1. 오른쪽에서 왼쪽으로, 두 번째 자리 숫자마다 값을 두 배로 만드세요. 곱한 값이 9보다 크면, 그 곱의 각 자리 숫자를 더하세요.
2. 모든 자리 숫자의 합을 구하세요.
3. 모든 자리 숫자의 합이 10의 배수이면 그 숫자는 유효하며, 그렇지 않으면 유효하지 않습니다.

검증 숫자가 추가되어 7992739871x 형태가 될 계좌 번호 "7992739871" 예를 가정하세요:

```markdown
Account number      7   9  9  2  7  3  9   8  7  1  x
Double every other  7  18  9  4  7  6  9  16  7  2  x
Sum 2-char digits   7   9  9  4  7  6  9   7  7  2  x
```

`pass` 문장을 변수 `sum_of_odd_digits`과 값 `0`로 바꾸세요.

# --hints--

`sum_of_odd_digits = 0`가 `verify_card_number` 함수 안에 있어야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*=\s*0/);
    }
})
```

`pass`가 `verify_card_number` 함수 안에 있으면 안 됩니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.notMatch(function_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    pass

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
