---
id: 656877f71bba2b97acedf9af
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

Python에는 문자열 조작에 도움이 되는 내장 클래스가 있습니다. 그중 하나가 `str` 클래스입니다. 이 클래스에는 `maketrans`라는 메서드가 있어 변환 표를 만들 수 있습니다. 이 표는 문자열 내 문자를 교체하는 데 사용할 수 있습니다:

```python
str.maketrans({'t': 'c', 'l': 'b'})
```

위 코드를 문자열에 호출하면 모든 `t` 문자를 `c`로, 모든 `l` 문자를 `b`로 교체합니다.

`card_translation`라는 변수를 만들고 모든 `-` 및 ` ` 문자를 빈 문자열로 교체하는 변환 표를 할당하세요.

# --hints--

`card_translation` 내에 `main` 변수를 만들어야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_translation\s*=/);
    }
})
```

`card_translation`에 `str.maketrans({'-': '', ' ': ''})` 값을 할당해야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        const allowedMatches = [
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|")-\1\s*:\s*('|")\2\s*,\s*('|") \3\s*:\s*('|")\4\s*\}\s*\)/,
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|") \1\s*:\s*('|")\2\s*,\s*('|")-\3\s*:\s*('|")\4\s*\}\s*\)/,
        ];
        const anyMatch = allowedMatches.some((match) => match.test(function_body));
        assert.isTrue(anyMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    
--fcc-editable-region--
```
