---
id: 65687db6a9fcf3c548a6c1b7
title: 28단계
challengeType: 20
dashedName: step-28
---

# --description--

알고리즘의 일부는 오른쪽에서부터 시작하여 두 번째 숫자마다 두 배로 만드는 것입니다. 숫자를 두 배로 했을 때 결과가 `10` 이상이면, 두 자릿수를 더하세요. 예를 들어, 숫자가 `6`이면 두 배로 해서 `12`가 됩니다. `1`과 `2`를 더해 `3`를 구하세요. 이는 _정수 나눗셈_을 사용해 첫 번째 숫자를 구하고 나머지 연산자(`%`)를 사용해 두 번째 숫자를 구하면 됩니다.

```python
my_number = 12
first_digit = my_number // 10
second_digit = my_number % 10
```

정수 나눗셈은 나눗셈의 몫을 구하며, 가장 가까운 정수로 내림합니다.

`if` 문장 안에서 `number`에 `number // 10` (정수 나눗셈) 결과와 `number`과 `10`의 나머지를 더한 값을 할당하세요.

# --hints--

`number = number // 10 + number % 10` 문장 안에 `if`가 있어야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getBlock("\n" + transformedCode, "if number >= 10");
        const { block_body } = verify_card_number;
        const allowedMatches = [
            /number *= *\(?\s*number *\/\/ *10\s*\)? *\+ *\(?\s*number *% *10/,
            /number *= *\(?\s*number *% *10\s*\)? *\+ *\(?\s*number *\/\/ *10/
        ];

        const someMatch = allowedMatches.some(allowedMatch => allowedMatch.test(block_body));
        assert.isTrue(someMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

--fcc-editable-region--
    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            print(number)
            
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
