---
id: 656878585631369a6b2d2191
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

번역을 정의하는 것만으로는 문자열이 번역되지 않습니다. `translate` 메서드를 번역할 문자열에 대해 번역 표를 인수로 하여 호출해야 합니다.

```py
my_string = "tamperlot"
translation_table = str.maketrans({'t': 'c', 'l': 'b'})
translated_string = my_string.translate(translation_table)
```

`translated_card_number`이라는 변수를 만들고 `translate`를 인수로 하여 `card_number`에 `card_translation` 메서드를 호출한 결과를 할당하세요.

# --hints--

`translated_card_number` 내에 `main` 변수를 만들어야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=/);
    }
})
```

`translated_card_number`에 `card_number.translate(card_translation)` 값을 할당해야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=\s*card_number\.translate\(\s*card_translation\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    
--fcc-editable-region--
```
