---
id: 65687a8253766cac4c99c57f
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

`verify_card_number` 함수 내에서 `odd_digits` 문자열의 모든 홀수 번째 숫자를 포함하는 변수 `card_number_reversed`을 만드세요.

# --hints--

`odd_digits = card_number_reversed[::2]` 함수 내에 `verify_card_number`이 있어야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /odd_digits\s*=\s*card_number_reversed\s*\[\s*:\s*:\s*2\s*\]/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
