---
id: 656879613a96aba21cbe80b9
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

`main` 함수 안에서 `verify_card_number` 함수를 호출하고 `translated_card_number` 변수를 인수로 전달하세요.

# --hints--

`verify_card_number(translated_card_number)` 함수 안에 `main`가 있어야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +verify_card_number\(\s*translated_card_number\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    pass

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

--fcc-editable-region--
main()
```
