---
id: 656879c6f35a85a59c06b3a7
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

문자열의 요소(문자)에 `[]` 인덱스 연산자를 사용해 접근한 적이 있습니다. `string[start:stop:step]` 인덱스 연산자를 사용해 문자열에서 문자 범위에 접근할 수도 있습니다:

```python
my_string = 'camperbot'
my_string[0:6] == 'camper' # True
my_string[0:6:3] == 'cp' # True
```

여기서 `start`는 시작 인덱스(포함), `stop`는 끝 인덱스(제외), `step`는 건너뛸 문자 수입니다. 지정하지 않으면 `step`가 기본값으로 1입니다.

`card_number_reversed`의 처음 4문자를 값으로 갖는 `card_number`라는 변수를 만드세요.

# --hints--

`card_number_reversed = card_number[0:4]` 함수 안에 `verify_card_number`가 있어야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /card_number_reversed\s*=\s*card_number\s*\[\s*0\s*:\s*4\s*\]/);
    }
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
