---
id: 6569b831a07d7154c793301b
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

매개변수 `verify_card_number`을 가진 함수 `card_number`를 정의하고, 함수가 아무 작업도 하지 않도록 `pass` 키워드를 사용하세요.

# --hints--

`verify_card_number`이라는 매개변수를 가진 `card_number`라는 함수가 있어야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        assert.exists(verify_card_number);
        const { function_parameters } = verify_card_number;
        assert.equal(function_parameters.trim(), "card_number");
    }
})
```

`pass` 함수 본문에 `verify_card_number` 키워드를 사용해야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        assert.exists(verify_card_number);
        const { function_body } = verify_card_number;
        assert.match(function_body, / +pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--

def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

main()
```
