---
id: 656b481ec4976439565a78b2
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

단계가 선택 사항인 것처럼, `0`에서 시작하거나 슬라이스 끝에서 끝나는 것도 선택 사항입니다:

```python
my_string = 'camperbot'
camperbot = my_string[::]
```

전체 `card_number` 문자열의 역순을 `card_number_reversed` 변수에 할당하세요.

# --hints--

`card_number_reversed = card_number[::-1]` 함수 안에 `verify_card_number`이 있어야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /card_number_reversed\s*=\s*card_number\s*\[\s*:\s*:\s*-1\s*\]/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[-1:-5:-1]
    print(card_number_reversed)
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
