---
id: 66bb3e20d3dc5b6d0a21f5dd
title: リアルタイムカウンターを作成する
challengeType: 25
dashedName: build-a-real-time-counter
demoType: onClick
---

# --description--

このラボでは、JavaScriptを使ってリアルタイムの文字カウンターを作成します。文字カウンターは`textarea`要素に入力された文字数を表示します。ユーザーが`textarea`に入力するたびにカウンターはリアルタイムで更新されます。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `textarea`の`id`を持つ`text-input`要素が必要です。
2. `p`の`id`を持つ`char-count`要素が必要です。最初はテキスト`Character Count: 0/50`を含んでいる必要があります。このプレースホルダーテキストは`textarea`要素内の現在の文字数に置き換えられます。
3. `#text-input`要素がテキスト`hello coder`を含むとき、`#char-count`要素はテキスト`"Character Count: 11/50"`を含む必要があります。
4. `#text-input`要素がテキスト`i am a programmer`を含むとき、`#char-count`要素はテキスト`"Character Count: 17/50"`を含む必要があります。
5. `#text-input`要素がテキスト`hello world`を含むとき、`#char-count`要素はテキスト`"Character Count: 11/50"`を含む必要があります。
6. `#text-input`要素がテキスト`I am learning a new language and it is called c++.`を含むとき、`#char-count`要素はテキスト`"Character Count: 50/50"`を含む必要があります。**注意:** `maxlength`属性を使えば正しい機能は実現できますが、このラボではJavaScriptのみの解決策が必要です。  
7. ユーザーは`50`要素に`textarea`文字を超えて入力できません。文字数が`50`に達したとき、余分な入力は自動的に切り詰められ、テキスト`Character Count: 50/50`が`red`に表示されます。

**注意:** スタイルシートとJavaScriptファイルをHTMLにリンクするのを忘れないでください。

# --hints--

`textarea`の`id`を持つ`text-input`要素が必要です。

```js
const textInput = document.querySelector('textarea');
assert.equal(textInput.id, 'text-input');
```

`p`の`id`を持つ`char-count`要素が必要です。

```js
assert.exists("p#char-count");
```

`#char-count`要素は最初にテキスト`Character Count: 0/50`を含む必要があります。

```js
const charCount = document.querySelector('p#char-count')?.innerText;
assert.equal(charCount, 'Character Count: 0/50');
```

`#text-input`要素がテキスト`hello coder`を含むとき、`#char-count`要素はテキスト`"Character Count: 11/50"`を含む必要があります。

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "hello coder";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(charCount.innerText.trim(), 'Character Count: 11/50');
```

文字数が`50`のとき、テキストは`red`に表示されます。

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "I am learning a new language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
const computedStyle = window.getComputedStyle(charCount);
assert.oneOf(computedStyle.color, [
  "red",
  "rgb(255, 0, 0)",
  "rgb(100%, 0%, 0%)",
  "rgba(255, 0, 0, 1)",
  "rgba(100%, 0%, 0%, 1)",
  "hsl(0, 100%, 50%)",
  "hsla(0, 100%, 50%, 1)"
]);
```

文字数が`50`以上の場合、余分な入力は切り詰められます。

```js
const textInput = document.getElementById('text-input');
textInput.value = "I am learning a new programming language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(textInput.value.length, 50);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real Time Counter</title>
</head>

<body>

</body>


</html>
```

```css

```

```js
    
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real-time Character Counter</title>
    <link rel="stylesheet" href="styles.css">

<body>
    <div class="container">
        <h1>Real-Time Character Counter</h1>
        <textarea id="text-input" placeholder="Type something..."></textarea>
        <p id="char-count">Character Count: 0/50</p>
    </div>

<script src="script.js"></script>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    font-family: Arial, sans-serif;
    background-color: #c0c1ec;
    margin: 0;
}

h1{
    background-color: #f2f48b;
    padding-left: 10px;
    padding-right: 10px;
}
.container {
    text-align: center;
}

textarea {
    width: 300px;
    height: 100px;
    font-size: 18px;
    padding: 10px;
    border: 1px solid #ccc;
    border-radius: 5px;
    resize: none;
    display: block;
    margin: 0 auto;
    border: 1px solid slategray;
    box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

textarea:focus {
    outline: none;
    border-color: #f2f48b;
}

#char-count {
    margin-top: 10px;
    font-size: 18px;
    color: darkslategray;
    font-weight: bold;
}
```

```js
document.getElementById("text-input").addEventListener("input", function () {
  let textInput = document.getElementById("text-input");
  let charCount = textInput.value.length;

  if (charCount > 50) {
    textInput.value = textInput.value.substring(0, 50);
    charCount = 50;
  }

  let charCountDisplay = document.getElementById("char-count");
  charCountDisplay.innerHTML = `Character Count: ${charCount}/50`;

  if (charCount === 50) {
    charCountDisplay.style.color = "red";
  } 

});
```
