---
id: 5dc17dc8f86c76b9248c6eb5
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

التعليق يسمح لك بترك رسائل دون التأثير على عرض المتصفح. كما يسمح لك بجعل الشفرة غير نشطة. يبدأ التعليق في HTML بـ `<!--`، ويحتوي على أي عدد من أسطر النص، وينتهي بـ `-->`.

إليك مثالًا على تعليق يحتوي على النص `TODO: Remove h1`:

```html
<!-- TODO: Remove h1 -->
```

أضف تعليقًا فوق عنصر `p` بهذا النص:

`TODO: Add link to cat photos`

# --hints--

يجب أن يبدأ تعليقك بـ `<!--`. أنت تفتقد حرفًا أو أكثر من الأحرف التي تحدد بداية التعليق.

```js
assert.match(code, /<!--/);
```

يجب أن ينتهي تعليقك بـ `-->`. أنت تفتقد حرفًا أو أكثر من الأحرف التي تحدد نهاية التعليق.

```js
assert.match(code, /-->/);
```

يجب ألا تحتوي الشفرة على أحرف تعليق إضافية للفتح أو الإغلاق. لديك `<!--` أو `-->` إضافية تظهر في المتصفح.

```js
const noSpaces = code.replace(/\s/g, '');
assert.isBelow(noSpaces.match(/<!--/g)?.length, 2)
assert.isBelow(noSpaces.match(/-->/g)?.length, 2);
```

يجب أن يكون تعليقك فوق عنصر `p`. لقد وضعتها بترتيب خاطئ.

```js
assert.match(
  code.replace(/\s/g, ''),
  /<!--(.*?)--><p>everyonelovescutecatsonline!<\/p>/i
);
```

يجب أن يحتوي تعليقك على النص `TODO: Add link to cat photos`.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
--fcc-editable-region--
    <p>Everyone loves cute cats online!</p>
  </body>
</html>
```
