---
id: 5dc17d3bf86c76b9248c6eb4
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

أنشئ عنصر `p` أسفل عنصر `h2` وأضف إليه النص التالي:

`Everyone loves cute cats online!`

# --hints--

يجب أن يحتوي عنصر `p` على وسم بداية. صيغة وسم البداية هي: `<elementName>`.

```js
assert.exists(document.querySelector('p'));
```

يجب أن تكون وسوم `p` بأحرف صغيرة. حسب الاصطلاح، تُكتب جميع وسوم HTML بأحرف صغيرة.

```js
assert.notMatch(code, /<\/?P>/);
```

يجب أن يحتوي عنصر `p` على وسم الإغلاق. يحتوي وسم الإغلاق على `/` مباشرة بعد الحرف `<`.

```js
assert.match(code, /<\/p\>/);
```

يجب أن يكون نص عنصر `p` هو `Everyone loves cute cats online!`، إما أنك حذفت النص أو به خطأ مطبعي.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  ?.innerText?.trim().replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

يجب أن يكون عنصر `p` أسفل عنصر `h2`. لقد وضعتهم بترتيب خاطئ.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H2'), collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
