---
id: 5dc174fcf86c76b9248c6eb2
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستواصل العمل مع عناصر HTML الأساسية مثل العناوين، الفقرات، والقوائم من خلال بناء تطبيق صور قطط.

ابدأ الورشة بإضافة عنصر `h1` يحتوي على النص `CatPhotoApp`.

# --hints--

يجب أن يكون النص `CatPhotoApp` موجودًا في الكود. قد ترغب في التحقق من تهجئتك.

```js
assert.match(code, /catphotoapp/i);
```

يجب أن يحتوي عنصر `h1` على وسم بداية. وسم البداية له هذه الصيغة النحوية: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

يجب أن تكون وسوم `h1` بأحرف صغيرة. حسب الاصطلاح، تُكتب جميع وسوم HTML بأحرف صغيرة.

```js
assert.notMatch(code, /<\/?H1>/);
```

يجب أن يحتوي عنصر `h1` على وسم إغلاق. وسم الإغلاق له هذه الصيغة النحوية: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

يجب أن يكون نص عنصر `h1` هو `CatPhotoApp`. إما أنك حذفت النص، أو به خطأ مطبعي، أو أنه ليس بين وسم البداية ووسم الإغلاق لعنصر `h1`.

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
