---
id: 5dc1798ff86c76b9248c6eb3
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

أسفل عنصر `h1`، أضف عنصر `h2` بهذا النص:

`Cat Photos`

# --hints--

يجب أن يحتوي عنصر `h1` على وسم بداية. وسم البداية له هذه الصيغة النحوية: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

يجب أن يحتوي عنصر `h1` على وسم إغلاق. وسم الإغلاق له هذه الصيغة النحوية: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

يجب أن يكون لديك عنصر `h1` واحد فقط. احذف العنصر الزائد.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

يجب أن يكون نص عنصر `h1` هو `CatPhotoApp`. إما أنك حذفت النص أو هناك خطأ مطبعي.

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```

يجب أن تكون وسوم `h2` بأحرف صغيرة. حسب الاصطلاح، تُكتب جميع وسوم HTML بأحرف صغيرة.

```js
assert.notMatch(code, /<\/?H2>/);
```

يجب أن يحتوي عنصر `h2` على وسم بداية. صيغ وسوم البداية هي: `<elementName>`.

```js
assert.exists(document.querySelector('h2'));
```

يجب أن يحتوي عنصر `h2` على وسم إغلاق. تحتوي وسوم الإغلاق على `/` مباشرة بعد الحرف `<`.

```js
assert.match(code, /<\/h2\>/);
```

يجب أن يكون نص عنصر `h2` هو `Cat Photos`. ضع فقط النص `Cat Photos` بين وسم البداية ووسم الإغلاق لعنصر `h2`.

```js
assert.equal(document.querySelector('h2')?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat photos');
```

يجب أن يكون عنصر `h2` أسفل عنصر `h1`. عنصر `h1` له أهمية أكبر ويجب أن يكون فوق عنصر `h2`.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H1'), collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
