---
id: 5dc23f9bf86c76b9248c6eba
title: خطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

يمكنك إضافة صور إلى موقعك باستخدام عنصر `img`. عناصر `img` لها وسم بداية بدون وسم إغلاق. العنصر الذي لا يحتوي على وسم إغلاق يُعرف باسم <dfn>العنصر الفارغ</dfn>.

أضف عنصر `img` أسفل عنصر `p`. في هذه المرحلة، لن تظهر أي صورة في المتصفح.

# --hints--

يجب أن يحتوي عنصر `img` على وسم بداية. أوسمة البداية لها هذه الصيغة النحوية: `<elementName>`.

```js
assert.exists(document.querySelector('img'));
```

يجب ألا يحتوي عنصر `img` على وسم إغلاق `</img>`.

```js
assert.notMatch(code, /<\/img\>/);
```

يجب أن يكون لديك عنصر `img` واحد فقط. احذف أي عناصر إضافية.

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

يجب أن يكون عنصر `img` أسفل عنصر `p`. لقد وضعتهم بالترتيب الخطأ.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
