---
id: 5dc24614f86c76b9248c6ebd
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

يمكنك الربط بصفحة أخرى باستخدام عنصر وسم الرابط (`a`).

إليك مثالًا يربط بـ `https://www.freecodecamp.org`:

```html
<a href="https://www.freecodecamp.org"></a>
```

أضف عنصر وسم الرابط بعد الفقرة يربط بـ `https://freecatphotoapp.com`. في هذه المرحلة، لن يظهر الرابط في المعاينة.

# --hints--

يجب أن يحتوي عنصر وسم الرابط (`a`) على وسم بداية. وسم البداية له هذه الصيغة النحوية: `<elementName>`.

```js
assert.exists(document.querySelector('a'));
```

يجب أن يحتوي عنصر وسم الرابط (`a`) على وسم إغلاق. وسم الإغلاق يحتوي على `/` مباشرة بعد الحرف `<`.

```js
assert.match(code, /<\/a\>/);
```

يجب أن يكون عنصر وسم الرابط (`a`) أسفل عنصر `p`. لقد وضعتهم بترتيب خاطئ.

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('A'));
```

عنصر وسم الرابط (`a`) لا يحتوي على خاصية `href`. تحقق من وجود مسافة بعد اسم وسم البداية و/أو وجود مسافات قبل أسماء جميع الخصائص.

```js
assert.isTrue(document.querySelector('a')?.hasAttribute('href'));
```

يجب أن يربط عنصر وسم الرابط (`a`) بـ `https://freecatphotoapp.com`. إما أنك حذفت عنوان URL أو هناك خطأ مطبعي.

```js
assert.equal(
  document.querySelector('a')?.getAttribute('href'),
    'https://freecatphotoapp.com'
);
```

على الرغم من أنك ضبطت خاصية `href` لعنصر وسم الرابط (`a`) على الرابط الصحيح، يُنصح دائمًا بوضع قيمة الخاصية بين علامات اقتباس.

```js
assert.notMatch(
  code, /\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
