---
id: 5dc24073f86c76b9248c6ebb
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

تُعد خاصيات HTML كلمات خاصة تُستخدم داخل وسم البداية لعنصر للتحكم في سلوك العنصر. تحدد الخاصية `src` في عنصر `img` عنوان URL للصورة (مكان وجود الصورة).

إليك مثالًا على عنصر `img` مع خاصية `src` تشير إلى شعار freeCodeCamp:

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

داخل عنصر `img` الموجود، أضف خاصية `src` مع هذا العنوان URL:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

يجب أن يحتوي شفرتك على عنصر `img`. ربما أزلت عنصر `img` أو لم تحط قيمة الخاصية `src` بعلامات اقتباس.

```js
assert.exists(document.querySelector('img'));
```

يجب أن يحتوي عنصر `img` على خاصية `src`. إما أنك أغفلت الخاصية أو بها خطأ مطبعي. تأكد من وجود مسافة بين اسم العنصر واسم الخاصية.

```js
assert.exists(document.querySelector('img')?.src);
```

يجب تعيين خاصية `src` لعنصر `img` إلى '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'. إما أنك أغفلت عنوان URL أو به خطأ مطبعي. حساسية حالة الأحرف في العنوان مهمة.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

رغم تعيين خاصية `src` لعنصر `img` إلى العنوان الصحيح، يُنصح دائمًا بحصر قيمة الخاصية بعلامات اقتباس.

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

يجب إغلاق عنصر `img` إما بـ `>` أو `/>`.

```js
assert.match(code, /<img\s+src\s*=\s*("|')https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```
