---
id: 5dfb6250eacea3f48c6300b2
title: الخطوة 24
challengeType: 0
dashedName: step-24
---

# --description--

بعد القائمة غير المرتبة، أضف صورة جديدة بقيمة خاصية `src` مضبوطة على:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

وقيمة خاصية `alt` الخاصة بها إلى:

`A slice of lasagna on a plate.`

# --hints--

يجب أن يكون هناك عنصر `img` مباشرة بعد وسم الإغلاق `</ul>`.

```js
assert.equal(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.nodeName, 'IMG');
```

الصورة الجديدة لا تحتوي على خاصية `alt`. تحقق من وجود مسافة بعد اسم وسم البداية و/أو وجود مسافات قبل أسماء جميع الخصائص.

```js
assert.isTrue(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.hasAttribute('alt'));
```

يجب أن تحتوي الصورة الجديدة على قيمة `alt` تساوي `A slice of lasagna on a plate.` تأكد من أن قيمة خاصية `alt` محاطة بعلامات اقتباس.

```js
assert.match(
  document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.getAttribute('alt')?.replace(/\s+/g, ' '), /^A slice of lasagna on a plate\.?$/i 
);
```

الصورة الجديدة لا تحتوي على خاصية `src`. تحقق من وجود مسافة بعد اسم وسم البداية و/أو وجود مسافات قبل أسماء جميع الخصائص.

```js
assert.isTrue(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.hasAttribute('src'));
```

يجب أن تحتوي الصورة الجديدة على قيمة `src` تساوي `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`. تأكد من أن قيمة خاصية `src` محاطة بعلامات اقتباس.

```js
assert.equal(
  document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.getAttribute('src'),
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

رغم أنك ضبطت خاصية `src` للصورة الجديدة على عنوان URL الصحيح، يُنصح دائمًا بإحاطة قيمة الخاصية بعلامات اقتباس.

```js
assert.notMatch(code, /\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
