---
id: 5ef9b03c81a63668521804e8
title: الخطوة 36
challengeType: 0
dashedName: step-36
---

# --description--

قم بتسكين عنصر `p` الذي يحتوي على النص `No Copyright - freeCodeCamp.org` داخل عنصر الـ `footer`.

# --hints--

لقد حذفت عنصر الـ `footer` أو أنه يفتقد وسم البداية أو وسم الإغلاق.

```js
assert.exists(document.querySelector('footer'));
assert.match(code, /<\/footer>/);
```

يجب أن يحتوي عنصر الـ `footer` على عنصر `p`. تأكد من إضافة وسم البداية ووسم الإغلاق لعنصر الـ `p`.

```js
assert.exists(document.querySelector('footer > p'));
```

يجب أن يحتوي عنصر الـ `footer` على وسم الإغلاق. وسم الإغلاق يحتوي على `/` مباشرة بعد حرف `<`.

```js
const pElemClosingTags = code.match(/<\/p\>/g);
assert.exists(pElemClosingTags);
assert.lengthOf(pElemClosingTags, 3);
```

يجب أن يكون نص عنصر الـ `p` هو `No Copyright - freeCodeCamp.org`. إما أنك حذفت النص أو به خطأ مطبعي، أو أنه ليس بين وسم البداية ووسم الإغلاق لعنصر الـ `p`.

```js
const extraSpacesRemoved = document.querySelectorAll('footer > p')[0]?.innerText?.trim().replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /^No Copyright - freeCodeCamp\.org$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      
--fcc-editable-region--      
    </footer>
  </body>
</html>
```
