---
id: 5ef9b03c81a63668521804ec
title: الخطوة 40
challengeType: 0
dashedName: step-40
---

# --description--

لاحظ أن محتويات الصفحة كلها مسكنة داخل عنصر `html`. عنصر `html` هو العنصر الجذري لصفحة HTML ويغلف كل المحتوى في الصفحة.

يمكنك أيضًا تحديد لغة الصفحة بإضافة الخاصية `lang` إلى عنصر `html`.

أضف الخاصية `lang` بالقيمة `en` إلى وسم البداية `html` لتحديد أن لغة الصفحة هي الإنجليزية.

# --hints--

لقد حذفت عنصر `html` أو أنه يفتقد وسم البداية أو وسم الإغلاق.

```js
assert.match(code, /\<html.*?\>/);
assert.match(code, /\<\/html\>/);
```

يجب أن يحتوي عنصر `html` على الخاصية `lang` بالقيمة `en`. ربما أغفلت الخاصية أو القيمة، أو بها خطأ مطبعي.

```js
const extraSpacesRemoved = code.replace(/\s+/g, " ");
assert.match(extraSpacesRemoved, /\<html\s+lang\s*\=\s*("?|'?)en\1\s*\>/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<html>
--fcc-editable-region--
  <head>
    <title>CatPhotoApp</title>
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
