---
id: 5ef9b03c81a63668521804d0
title: خطوة 27
challengeType: 0
dashedName: step-27
---

# --description--

لتأكيد كلمة أو عبارة معينة، يمكنك استخدام العنصر `em`.

أكد كلمة `love` في عنصر `figcaption` عن طريق تغليفها بعنصر تأكيد `em`.

# --hints--

يجب أن يحتوي عنصر التأكيد `em` على وسم بداية. صيغتها النحوية هي: `<elementName>`.

```js
assert.exists(document.querySelector('em'));
```

يجب أن يحتوي عنصر التأكيد `em` على وسم إغلاق. وسم الإغلاق يحتوي على `/` مباشرة بعد الحرف `<`.

```js
assert.match(code, /<\/em\>/);
```

لقد حذفت إما عنصر `figcaption` أو أنه يفتقد وسم بداية أو وسم إغلاق.

```js
assert.exists(document.querySelector('figcaption'));
assert.match(code, /<\/figcaption\>/);
```

يجب أن يحيط عنصر التأكيد `em` بالنص `love`. إما أنك حذفت النص أو به خطأ مطبعي.

```js
assert.equal(
  document.querySelector('figcaption > em')?.innerText?.trim().toLowerCase(), 'love'
);
```

يجب أن يكون نص `figcaption` هو `Cats love lasagna`. تحقق من الأخطاء المطبعية وأن المسافات اللازمة موجودة حول وسم البداية ووسم الإغلاق لعنصر `em`.

```js
assert.match(
  document
    .querySelector('figcaption')
    ?.innerText?.trim().replace(/\s+/gi, ' '),
    /^cats love lasagna\.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```
