---
id: 5ef9b03c81a63668521804e7
title: الخطوة 35
challengeType: 0
dashedName: step-35
---

# --description--

يُستخدم عنصر `footer` لتعريف تذييل لوثيقة أو قسم. عادةً يحتوي التذييل على معلومات عن مؤلف الوثيقة، بيانات حقوق النشر، روابط لشروط الاستخدام، معلومات الاتصال، والمزيد.

بعد عنصر `main`، أضف عنصر `footer`.

# --hints--

لقد حذفت عنصر `main` أو أنه يفتقد وسم البداية أو وسم الإغلاق.

```js
assert.exists(document.querySelector('main'));
assert.match(code, /<\/main>/);
```

يجب أن يحتوي عنصر `footer` على وسم بداية. تحتوي وسوم البداية على الصيغة النحوية التالية: `<elementName>`.

```js
assert.exists(document.querySelector('footer'));
```

يجب أن يحتوي عنصر `footer` على وسم إغلاق. تحتوي وسوم الإغلاق على `/` مباشرة بعد الحرف `<`.

```js
assert.match(code, /<\/footer\>/);
```

يجب أن يكون عنصر `footer` أسفل وسم إغلاق عنصر `main`. لقد وضعته في مكان آخر.

```js
assert.equal(document.querySelector('main')?.nextElementSibling.nodeName, 'FOOTER');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
