---
id: 5ef9b03c81a63668521804e9
title: الخطوة 37
challengeType: 0
dashedName: step-37
---

# --description--

حوّل النص الموجود `freeCodeCamp.org` إلى رابط عن طريق تضمينه في عنصر وسم الرابط (`a`). يجب تعيين خاصية `href` إلى `https://www.freecodecamp.org`.

# --hints--

يجب أن يكون وسم الرابط (`a`) مسكنًا داخل وسم الفقرة (`p`) داخل التذييل (`footer`). تأكد من إضافة وسم بداية ووسم إغلاق لوسم الرابط (`a`).

```js
assert.isNotEmpty(document.querySelectorAll('footer > p > a'));
```

يجب أن يحتوي عنصر وسم الرابط (`a`) على وسم إغلاق. وسم الإغلاق يحتوي على `/` مباشرة بعد الحرف `<`.

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert.exists(aElemClosingTags);
assert.lengthOf(aElemClosingTags, 4);
```

يجب أن يحتوي عنصر وسم الرابط (`a`) على خاصية `href` بالقيمة `https://www.freecodecamp.org`. قد تكون قد أغفلت الخاصية/القيمة أو بها خطأ مطبعي.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(nestedAnchor?.getAttribute('href').toLowerCase(), 'https://www.freecodecamp.org');
```

يجب أن يكون نص الرابط هو `freeCodeCamp.org`. قد تكون أغفلت النص أو به خطأ مطبعي.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(
  nestedAnchor?.innerText?.trim().toLowerCase(),
    'freecodecamp.org'
);
```

بعد التسكين لعنصر وسم الرابط (`a`)، يجب أن يكون المحتوى الوحيد الظاهر لعنصر `p` في المتصفح هو `No Copyright - freeCodeCamp.org`. تحقق جيدًا من النص أو المسافات أو علامات الترقيم لكل من عنصر `p` وعنصر وسم الرابط المسكن.

```js
const pText = document.querySelectorAll('footer > p')[0]?.innerText?.trim().toLowerCase().replace(/\s+/g, ' ');
assert.match(pText, /^no copyright - freecodecamp\.org$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```
