---
id: 5f07c98cdb9413cbd4b16750
title: خطوة 19
challengeType: 0
dashedName: step-19
---

# --description--

حان الوقت لإضافة قسم جديد. أضف عنصر `section` ثانيًا أسفل عنصر `section` الموجود.

# --hints--

يجب أن يحتوي عنصر `section` على وسم بداية. وسم البداية له هذه الصيغة النحوية: `<elementName>`.

```js
assert.isAtLeast(document.querySelectorAll('section').length, 2);
```

يجب أن تضيف وسم بداية `section` واحد فقط. احذف أي إضافات زائدة.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
```

يجب أن يحتوي عنصر `section` على وسم إغلاق. وسم الإغلاق يحتوي على `/` مباشرة بعد حرف `<`.

```js
assert.isAtLeast(code.match(/<\/section>/g)?.length, 2);
```

يجب أن تضيف وسم إغلاق `section` واحد فقط. احذف أي إضافات زائدة.

```js
assert.lengthOf(code.match(/<\/section>/g), 2);
```

يجب ألا يكون عنصر `section` الثاني مسكنًا داخل عنصر `section` الأول.

```js
const el = document.querySelector('main > section');
assert.isNotNull(el);
const childrenOf1stSection = [
  ...el.children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.isEmpty(foundElems);
```

يجب أن يكون كلا العنصرين `section` بين وسم البداية ووسم الإغلاق لعنصر `main`.

```js
const el = document.querySelector('main');
assert.isNotNull(el);
const childrenOfMain = [...el.children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.lengthOf(foundElems, 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
