---
id: 671141feba228a35cefba82d
title: الخطوة 14
challengeType: 0
dashedName: step-14
---

# --description--

حوّل النص الموجود `cute cats` إلى عنصر وسم الرابط الذي يربط إلى:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`

# --hints--

يجب أن يكون هناك عنصر وسم رابط جديد داخل أول عنصر `p`.

```js
assert.exists(document.querySelector('p')?.querySelector('a'));
```

يجب أن يكون نص وسم الرابط هو `cute cats`.

```js
assert.strictEqual(document.querySelector('a')?.innerText?.trim(), "cute cats");
```

يجب أن تكون خاصية `href` لعنصر وسم الرابط الجديد هي `https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`.

```js
assert.strictEqual(document.querySelector('a')?.href, "https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg");
```

يجب أن يكون هناك نسخة واحدة فقط من العبارة `cute cats` في شفرتك.

```js
assert.strictEqual(code.indexOf("cute cats"),code.lastIndexOf("cute cats")); 
```

يجب أن يظل نص عنصر `p` هو `Everyone loves cute cats online!`.

```js
assert.strictEqual(document.querySelector('p')?.innerText?.trim(), "Everyone loves cute cats online!");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
