---
id: 671b6e873249bb35c9debfcf
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

أضف الكلمات `See more ` قبل عنصر وسم الرابط و` in our gallery` بعد عنصر وسم الرابط.

# --hints--

يجب استخدام عنصر وسم الرابط الموجود، لا تنشئ عنصرًا جديدًا.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

يجب ألا يتغير النص داخل عنصر وسم الرابط، يجب أن يكون `cat photos`.

```js
assert.strictEqual(document.querySelector('a')?.innerText?.replace(/\s+/g, ' ').trim(), "cat photos");
```

يجب أن يحتوي شفرتك على `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery`.

```js
assert.match(code, /See more\s+<a href="https:\/\/freecatphotoapp\.com">cat photos<\/a>\s+in our gallery/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
