---
id: 69a0179a85b00b6127e0af45
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

Всередині циклу `while` оновіть `current` до `current.next`, щоб цикл рухався вперед по списку по одному вузлу за раз і зупинявся, коли досягне вузла з властивістю `next` зі значенням `null`, тобто кінця списку.

# --hints--

Ви повинні оновити значення вашої змінної `current` до `current.next` всередині циклу `while`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            --fcc-editable-region--

            --fcc-editable-region--
        }
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
