---
id: 69a0102e761a245bb1cbd8e7
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

На цьому етапі слід протестувати зв’язаний список для кращого розуміння.

Створіть змінну `myList` і присвойте їй значення `initList()`. Під нею виведіть у консоль виклик `isEmpty()` з `myList` як параметром. Потім викличте `add` з `myList` і числом на ваш вибір, після чого виведіть у консоль `myList`. Під цим знову виведіть у консоль виклик `isEmpty()` з `myList` як параметром.

# --hints--

Ви повинні створити змінну `myList`, присвоївши їй `initList()`.

```js
assert.deepEqual(myList, initList());
```

Виведіть у консоль виклик `isEmpty(myList)`.

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedCode, 'console.log(isEmpty(myList))');
```

Викличте `add` з `myList` і числом на ваш вибір як параметрами.

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
assert.match(strippedCode, /add\(myList,\d+\)/);
```

Виведіть у консоль вашу змінну `myList`.

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedCode, 'console.log(myList)');
```

Знову виведіть у консоль `isEmpty(myList)`.

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
const regexMatch = strippedCode.matchAll(/console\.log\(isEmpty\(myList\)\)/g);
assert.isAtLeast([...regexMatch].length, 2);
```


# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
