---
id: 69a00f673f171e29159e3ad8
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Цей вузол має мати дві властивості: елемент для додавання та вказівник на наступний вузол.

Усередині функції add створіть об’єкт з іменем `node`, він повинен мати властивість з іменем `element` і властивість `next`, встановлену в `null`, яка є вказівником на наступний елемент списку.

# --hints--

Ви маєте створити об’єкт з іменем `node` усередині вашої функції `add`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const { node } = add?.variables ?? {};
assert.match(node?.value.toString(), /^\{[\s\S]*\}/);
assert.exists(node);
```

Ваш об’єкт `node` має мати властивість з іменем `element`, встановлену в `element`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const { node } = add?.variables ?? {};
const { element } = node?.objectProps ?? {};
assert.isTrue(element?.value.matches('element'));
```

Ваш об’єкт `node` має мати властивість з іменем `next`, встановлену в `null`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const { node } = add?.variables ?? {};
const { next } = node?.objectProps ?? {};
assert.isTrue(next?.value.matches('null'));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}
```
