---
id: 69a00eaa43458e618331873c
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Тепер ви працюватимете над функцією, яка додасть вузол до списку.

Створіть функцію `add` з параметрами `list` та `element`.

# --hints--

Ви повинні створити функцію `add`.

```js
assert.isFunction(add)
```

Ваша функція `add` має мати параметри `list` та `element`.

```js
const addParams = __helpers.getFunctionParams(add.toString());
const addParamsNames = addParams.map(param => param.name)
assert.include(addParamsNames, "list");
assert.include(addParamsNames, "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
