---
id: 69a04466917c66b132323a1b
title: Крок 18
challengeType: 1
dashedName: step-18
---

# --description--

Якщо `previous` не є `null`, це означає, що елемент, який потрібно видалити, знаходиться не на початку, а десь усередині списку, тому поточний вузол слід обійти, зв’язавши попередній вузол безпосередньо з наступним.

Для цього створіть ще одну `if` інструкцію, яка перевіряє, чи попередній строго не є `null`. Якщо це так, встановіть `previous.next` у `current.next`.

# --hints--

Ви повинні створити ще одну `if` інструкцію з умовою `previous !== null`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*/);
```

Ви повинні встановити `previous.next` у `current.next`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*previous\.next\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }
--fcc-editable-region--
    
--fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
