---
id: 69a00cd4ec513345d40e95ec
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Наступне, над чим ви працюватимете, — це функція, яка дозволяє перевірити, чи є список порожнім.

Для цього створіть функцію з назвою `isEmpty`, яка приймає параметр `list`. Усередині функції поверніть перевірку, чи властивість `length` об’єкта `list` строго дорівнює `0`.

# --hints--

Ви повинні створити `isEmpty` функцію.

```js
assert.isFunction(isEmpty)
```

Ваша `isEmpty` функція має мати параметр `list`.

```js
const isEmptyParams = __helpers.getFunctionParams(isEmpty.toString());
assert.deepEqual(isEmptyParams[0].name, "list");
```

Ваша `isEmpty` функція має повертати `list.length === 0`.

```js
const emptyList = { head: null, length: 0 };
const nonEmptyList = { head: {}, length: 1 };

assert.isTrue(isEmpty(emptyList));
assert.isFalse(isEmpty(nonEmptyList));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
