---
id: 69a04088e600bd8074a9852f
title: Крок 17
challengeType: 1
dashedName: step-17
---

# --description--

Якщо `current` дорівнює `null`, це означає, що це кінець списку, тому функція має припинити виконання.

Створіть інструкцію `if`, яка перевіряє, чи `current` строго дорівнює `null`. Усередині інструкції `if` використайте `return` для виходу з функції.

# --hints--

Ви маєте створити інструкцію `if` з умовою `current === null`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(current\s*===\s*null\s*\)\s*\{\s*/);
```

Ви маєте використати ключове слово `return` всередині вашої інструкції `if`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(current\s*===\s*null\s*\)\s*\{\s*return;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
