---
id: 69a00a051b9a18694d00c16b
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви реалізуєте зв’язаний список.

Для початку створіть функцію з іменем `initList`. Усередині неї поверніть об’єкт із властивістю `head`, встановленою в `null`, та властивістю `length` зі значенням `0`.

Це функція, яку ви викликатимете для ініціалізації зв’язаного списку. `head` представляє перший елемент у списку, спочатку він дорівнює `null`, оскільки в списку немає елементів. `length` відстежує кількість вузлів (елементів) у списку, яку ви збільшуватимете або зменшуватимете при додаванні чи видаленні вузла зі списку.

# --hints--

Ви повинні створити функцію `initList`.

```js
assert.isFunction(initList)
```

Ваша функція `initList` має повертати об’єкт.

```js
assert.isObject(initList());
```

Ваша функція `initList` має повертати об’єкт із властивістю `head`, встановленою в `null`.

```js
assert.propertyVal(initList(), "head", null);
```

Ваша функція `initList` має повертати об’єкт із властивістю `length`, встановленою в `0`.

```js
assert.propertyVal(initList(), "length", 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
