---
id: 69a013a86c97dac3784b2636
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Тепер у консолі має з’явитися щось на кшталт `{ head: { element: 42, next: null }, length: 1 }`. 

Однак це не той спосіб, яким слід обробляти випадок, коли доданий вузол є першим елементом у списку, тобто головою.

Щоб зробити це правильно, використайте оператор if, щоб перевірити, чи порожній список, викликавши функцію `isEmpty` з `list`. Якщо це правда, тоді перемістіть свій рядок `list.head = node` у тіло інструкції `if`. 

# --hints--

У вас має бути інструкція `if` з умовою `isEmpty(list)`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*/);
```

Ви маєте перемістити свій рядок `list.head = node` в інструкцію `if`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*list\.head\s*=\s*node;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

--fcc-editable-region--
    list.head = node;
--fcc-editable-region--
    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
