---
id: 69a0102e761a245bb1cbd8e7
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

На цьому етапі слід протестувати зв’язаний список для кращого розуміння.

Створіть змінну `myList` і присвойте їй значення `initList()`. Під нею виведіть у консоль виклик `isEmpty()` з `myList` як параметром. Потім викличте `add` з `myList` і числом на ваш вибір, після чого виведіть у консоль `myList`. Під цим знову виведіть у консоль виклик `isEmpty()` з `myList` як параметром.

# --hints--

Ви повинні створити змінну `myList`, присвоївши їй `initList()`.

```js
assert.deepEqual(myList, initList());
```

Виведіть у консоль виклик `isEmpty(myList)`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(isEmpty\(myList\)\);?/)
```

Викличте `add` з `myList` і числом на ваш вибір як параметрами.

```js
assert.match(__helpers.removeJSComments(code), /add\s*\(\s*myList\s*,\s*\d+\s*\)/);
```

Виведіть у консоль вашу змінну `myList`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?/)
```

Знову виведіть у консоль `isEmpty(myList)`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?\s*console\.log\(isEmpty\(myList\)\);?/)
```


# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
