---
id: 69a04712866a37ea8f38b6bd
title: Крок 21
challengeType: 1
dashedName: step-21
---

# --description--

Тепер ви маєте протестувати функцію `remove`.

Видаліть будь-яке число зі списку, викликавши функцію `remove` з `myList` та числом, яке потрібно видалити, а потім виведіть `myList` у консоль за допомогою `console.log(JSON.stringify(myList, null, 2))`, щоб побачити повний оновлений список.

На цьому ваше практичне заняття зі зв’язаними списками завершене!

# --hints--

Ви маєте видалити будь-яке з існуючих чисел, використовуючи функцію `remove`.

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
const removalTargets = ['42', '43', '44'];
const removedOne = removalTargets.some(n => strippedCode.includes(`remove(myList,${n})`));
assert.isTrue(removedOne);
```

Ви маєте знову вивести `myList` у консоль, цього разу використовуючи `JSON.stringify(myList, null, 2)`.

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedCode, 'console.log(JSON.stringify(myList,null,2))');
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }

    if(previous !== null) {
        previous.next = current.next;
    } else {
        list.head = current.next;
    }

    list.length--;
    
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList);
console.log(isEmpty(myList));
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }

    if(previous !== null) {
        previous.next = current.next;
    } else {
        list.head = current.next;
    }

    list.length--;
    
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList);
console.log(isEmpty(myList));
remove(myList, 43);
console.log(JSON.stringify(myList, null, 2));
```
