---
id: 69a016ac6edea7d8c8020c98
title: Крок 9
challengeType: 1
dashedName: step-9
---

# --description--

Тепер створіть цикл `while` з умовою `current.next !== null`. Це означає, що цикл має продовжуватися, доки поточний вузол має наступний вузол.

# --hints--

Ви маєте створити цикл `while` з умовою `current.next !== null`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /while\s+\(current\.next\s*!==\s*null\)\s*\{\s*/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        --fcc-editable-region--

        --fcc-editable-region--
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
