---
id: 69a012c6224fac85832247dd
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Бачите результати `true`, `{ head: null, length: 0 }` і `true` у консолі? Це відбувається, тому що ви не призначили нічого списку `myList`, і не збільшили властивість `length`.

Призначте `node` для `list.head`, а під цим збільшіть властивість `length` об’єкта `list`.

# --hints--

Ви повинні призначити `node` для `list.head`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'list.head=node');
```

Ви повинні збільшити властивість `length` об’єкта `list`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'list.length++');
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    --fcc-editable-region--

    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
console.log(myList);
console.log(isEmpty(myList));
```
