---
id: 69a00f673f171e29159e3ad8
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Цей вузол має мати дві властивості: елемент для додавання та вказівник на наступний вузол.

Усередині функції add створіть об’єкт з іменем `node`, він повинен мати властивість з іменем `element` і властивість `next`, встановлену в `null`, яка є вказівником на наступний елемент списку.

# --hints--

Ви маєте створити об’єкт з іменем `node` усередині вашої функції `add`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /const|var|let\s+node\s*=\s*{/);
```

Ваш об’єкт `node` має мати властивість з іменем `element`, встановлену в `element`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /element\s*:\s*element/);
```

Ваш об’єкт `node` має мати властивість з іменем `next`, встановлену в `null`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /next\s*:\s*null/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}
```
