---
id: 69a046a1a74385bb4be1bb88
title: Крок 20
challengeType: 1
dashedName: step-20
---

# --description--

Тепер потрібно зменшити `length` у функції `remove`, щоб відстежувати фактичну кількість елементів, що залишилися в списку після видалення одного.

# --hints--

Ви повинні зменшити властивість `length` об’єкта `list` за допомогою оператора декременту (`--`).

```js
assert.match(__helpers.removeJSComments(remove.toString()), /list.length\-\-/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }

    if(previous !== null) {
        previous.next = current.next;
    } else {
        list.head = current.next;
    }

    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
