---
id: 69a014bd9b1730e94136279f
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Якщо список не порожній, потрібно знайти його кінець, додати новий елемент туди і зробити так, щоб він вказував на `null`. Це означає, що потрібно пройтися циклом по списку, починаючи з `head`.

Щоб почати, відкрийте умову `else`. Всередині неї використайте `let`, щоб створити змінну `current` і встановити її в `list.head`.

# --hints--

У вас має бути умова `else`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{/);
```

Ви маєте створити змінну `current`, встановлену в `list.head`, всередині умови `else`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{\s*let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    --fcc-editable-region--
    } 
    --fcc-editable-region--

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
