---
id: 69a03aaa1e6fe8152788ce6a
title: Крок 14
challengeType: 1
dashedName: step-14
---

# --description--

Всередині функції `remove` використайте `let` для створення змінної множини `previous`, встановленої в `null`, а також `let` для створення змінної множини `current`, встановленої в `list.head`. Саме ці змінні множини ви будете використовувати в циклі, який проходить по списку.

`current` відстежує вузол, який наразі перевіряється у списку, починаючи з голови і рухаючись вперед по одному вузлу за раз. `previous`, навпаки, відстежує вузол безпосередньо перед поточним.

Це важливо при видаленні вузла, оскільки потрібно оновити значення next попереднього вузла, щоб воно вказувало на вузол після того, що видаляється, щоб зберегти зв’язок списку.

# --hints--

Ви повинні використати `let` для створення змінної множини `previous`, встановленої в `null`, всередині вашої функції `remove`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+previous\s*=\s*null;?/);
```

Ви повинні використати `let` для створення змінної множини `current`, встановленої в `list.head`, всередині вашої функції `remove`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
