---
id: 687ad7be1ff45032ccf1d92f
title: Einen Theme-Switcher erstellen
challengeType: 25
dashedName: lab-theme-switcher
demoType: onClick
---

# --description--

In diesem Lab bauen Sie eine App, die zwischen verschiedenen Themes wechselt. Wenn Sie zu einem anderen Theme wechseln, ändert sich die Hintergrundfarbe und eine andere Nachricht wird angezeigt.

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bestehen Sie alle Tests, um das Lab abzuschließen.

**User Stories:**

1. Sie sollten ein `button`-Element mit dem Text `Switch Theme` haben.
2. Ihr `button`-Element sollte die folgenden Attribute haben:
   - Ein `id`-Attribut mit dem Wert `"theme-switcher-button"`.
   - Ein `aria-haspopup`-Attribut mit dem Wert `"true"`.
   - Ein `aria-expanded`-Attribut mit dem Wert `"false"`.
   - Ein `aria-controls`-Attribut mit dem Wert `"theme-dropdown"`.
3. Sie sollten ein `ul`-Element mit den folgenden Attributen haben:
   - Ein `id`-Attribut mit dem Wert `"theme-dropdown"`.
   - Ein `role`-Attribut mit dem Wert `"menu"`.
   - Ein `aria-labelledby`-Attribut mit dem Wert `"theme-switcher-button"`.
   - Ein `hidden`-Attribut.
4. Ihr `ul`-Element sollte mindestens zwei `li`-Elemente mit einem `role`-Attribut haben, das auf `"menuitem"` gesetzt ist, und Text Ihrer Wahl, der ein anderes Theme darstellt.
5. Jedes Ihrer `li`-Elemente sollte ein `id`-Attribut haben, das mit `theme-` beginnt und mit dem Theme endet, das Sie für den Text des `li`-Elements festgelegt haben. Zum Beispiel, wenn eines Ihrer Themes `Light` ist, sollte Ihr `id` `theme-light` sein.
6. Sie sollten ein Element mit einem `aria-live`-Attribut haben, das auf `"polite"` gesetzt ist.
7. Sie sollten ein `themes`-Array mit mindestens zwei Objekten haben, die jeweils eine `name`- und `message`-Eigenschaft enthalten. Die `name` stellt ein anderes Theme dar und die `message` wird angezeigt, wenn das Theme wechselt. Sie können `name`- und `message`-Werte Ihrer Wahl verwenden, aber die `name`-Werte sollten mit dem Textinhalt Ihrer `li`-Elemente übereinstimmen.
8. Wenn ein Benutzer auf das `#theme-switcher-button` klickt, sollte das `#theme-dropdown`-Element die verfügbaren Themes anzeigen.
9. Wenn das `#theme-dropdown`-Element angezeigt wird, sollte das `aria-expanded`-Attribut für das `"true"`-Element auf `button` gesetzt werden.
10. Wenn der Benutzer auf das `#theme-switcher-button` klickt, während das `#theme-dropdown`-Element angezeigt wird, sollte das Menü geschlossen werden, das `hidden`-Attribut hinzugefügt und das `aria-expanded`-Attribut auf `"false"` gesetzt werden.
11. Wenn ein Benutzer auf das `#theme-switcher-button` klickt und ein Theme auswählt, sollte die entsprechende `theme-<name>`-Klasse zum `body`-Element hinzugefügt werden und die zugehörige Theme-Nachricht, die Sie im `themes`-Array festgelegt haben, sollte im `aria-live="polite"`-Element angezeigt werden.
12. Wenn ein Benutzer auf das `#theme-switcher-button` klickt und ein anderes Theme als das aktuelle auswählt, sollten die `body`-Theme-Klasse und die Statusnachricht aktualisiert werden, um der neuen Auswahl zu entsprechen.

# --before-all--

```js
const resetThemeSwitch = () => {
  const themeDropdownBtn = document.getElementById("theme-switcher-button");
  const themeDropdownEl = document.getElementById("theme-dropdown");

  // Reset attributes due to earlier button clicks interfering with these tests
  themeDropdownEl.setAttribute("hidden", "");
  themeDropdownBtn.setAttribute("aria-expanded", "false");
}
```

# --hints--

Sie sollten ein `button`-Element haben.

```js
assert.exists(document.querySelector("button"));
```

Ihr `button`-Element sollte den Text `Switch Theme` haben.

```js
const btn = document.querySelector("button");
assert.equal(btn?.textContent.trim(), "Switch Theme");
```

Ihr `button`-Element sollte ein `id`-Attribut mit dem Wert `"theme-switcher-button"` haben.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("id"), "theme-switcher-button");
```

Ihr `button`-Element sollte ein `aria-haspopup`-Attribut mit dem Wert `"true"` haben.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-haspopup"), "true");
```

Ihr `button`-Element sollte ein `aria-expanded`-Attribut mit dem Wert `"false"` haben.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-expanded"), "false");
```

Ihr `button`-Element sollte ein `aria-controls`-Attribut mit dem Wert `"theme-dropdown"` haben.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-controls"), "theme-dropdown");
```

Sie sollten ein `ul`-Element haben.

```js
assert.exists(document.querySelector("ul"));
```

Ihr `ul`-Element sollte ein `id`-Attribut mit dem Wert `"theme-dropdown"` haben.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("id"), "theme-dropdown");
```

Ihr `ul`-Element sollte ein `role`-Attribut mit dem Wert `"menu"` haben.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("role"), "menu");
```

Ihr `ul`-Element sollte ein `aria-labelledby`-Attribut mit dem Wert `"theme-switcher-button"` haben.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("aria-labelledby"), "theme-switcher-button");
```

Ihr `ul`-Element sollte ein `hidden`-Attribut haben.

```js
const ul = document.querySelector("ul");
assert.isTrue(ul?.hasAttribute("hidden"));
```

Ihr `ul`-Element sollte mindestens zwei `li`-Elemente haben.

```js
const listItems = document.querySelectorAll("ul li");
assert.isAtLeast(listItems?.length, 2);
```

Jedes Ihrer `li`-Elemente sollte ein `role`-Attribut mit dem Wert `"menuitem"` haben.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  assert.equal(li?.getAttribute("role"), "menuitem");
})
```

Jedes Ihrer `li`-Elemente sollte einen Text enthalten, der ein Theme darstellt.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  const text = li?.textContent.trim();
  assert.isTrue(text.length > 0);
})
```

Jedes Ihrer `li`-Elemente sollte ein `id`-Attribut haben, das mit `theme-` beginnt und mit dem entsprechenden Theme für dieses Element endet. Stellen Sie sicher, dass Ihr `id`-Wert komplett in Kleinbuchstaben ist.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
const listTitles = [...listItems].map(li => li?.textContent.trim().toLowerCase());

listItems.forEach((li, index) =>{
  assert.equal(li.getAttribute("id"), `theme-${listTitles[index]}`);
})
```

Sie sollten ein Element mit einem `aria-live`-Attribut haben, das auf `"polite"` gesetzt ist.

```js
assert.exists(document.querySelector('[aria-live="polite"]'));
```

Sie sollten ein `themes`-Array haben.

```js
assert.isArray(themes);
```

Ihr `themes`-Array sollte mindestens zwei Objekte enthalten.

```js
assert.isAtLeast(themes?.length, 2);
themes.forEach(theme => {
  assert.isObject(theme);
});
```

Ihr `themes`-Array sollte mindestens zwei Objekte enthalten, die jeweils eine `name`- und `message`-Eigenschaft haben.

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.property(theme, "name");
  assert.property(theme, "message");
})
```

Ihre `name`- und `message`-Eigenschaftswerte sollten Strings sein, die nicht leer sind.

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.isString(theme["name"]);
  assert.isNotEmpty(theme["name"]);
  assert.isString(theme["message"]);
  assert.isNotEmpty(theme["message"]);
})
```

Jeder Ihrer `name`-Eigenschaftswerte sollte mit dem Textinhalt eines Ihrer `li`-Elemente übereinstimmen. Stellen Sie sicher, dass diese Werte komplett in Kleinbuchstaben sind.

```js
const listItems = [...document.querySelectorAll("ul li")]?.map(li => li.textContent.toLowerCase().trim());
const namesArr = themes?.map(({name}) => name);

assert.equal(listItems.length, namesArr.length);
assert.isTrue(namesArr?.every(name => listItems.includes(name)));
```

Wenn ein Benutzer auf das `#theme-switcher-button` klickt, um die Theme-Optionen anzuzeigen, sollte das `hidden`-Attribut vom `#theme-dropdown`-Element entfernt werden.

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
```

Wenn das `#theme-dropdown`-Element angezeigt wird, sollte das `aria-expanded`-Attribut für das `"true"`-Element auf `button` gesetzt werden.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
```

Wenn der Benutzer auf das `#theme-switcher-button` klickt, während das `#theme-dropdown`-Element angezeigt wird, sollte das Dropdown-Menü auf `hidden` zurückgesetzt werden.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
```

Wenn der Benutzer auf das `#theme-switcher-button` klickt, während das `#theme-dropdown`-Element angezeigt wird, sollte das `aria-expanded`-Attribut auf `"false"` gesetzt werden.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "false");
```

Wenn ein Benutzer auf das `#theme-switcher-button` klickt und ein Theme auswählt, sollte die entsprechende `theme-<name>`-Klasse zum `body`-Element hinzugefügt werden.

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("#theme-dropdown li");

themeDropdownBtn.dispatchEvent(new Event("click"));

const clickedItem = listItems[1];
clickedItem?.dispatchEvent(new Event("click", {bubbles: true}));

const expectedClass = clickedItem?.id;
const actualClass = document.body.className;

assert.equal(actualClass, expectedClass);
```

Wenn ein Benutzer auf das `#theme-switcher-button` klickt und ein Theme auswählt, sollte eine `message` angezeigt werden, die mit dem ausgewählten `theme` aus dem `themes`-Array zusammenhängt, und im `aria-live="polite"`-Element dargestellt werden.


```js
const liveRegion = document.querySelector('[aria-live="polite"]');
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("ul li");

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[1]?.dispatchEvent(new Event("click", {bubbles: true}));

const selectedThemeName = listItems[1]?.textContent.toLowerCase().trim();
const selectedTheme = themes?.find(theme => theme.name === selectedThemeName);

assert.equal(liveRegion?.textContent, selectedTheme?.message);
```

Wenn ein Benutzer auf das `#theme-switcher-button` klickt und zwei verschiedene Themes auswählt, sollten die `body`-Klasse und die Statusnachricht aktualisiert werden, um jedem ausgewählten Theme zu entsprechen.

```js
const liveRegion = document.querySelector('[aria-live="polite"]');
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("#theme-dropdown li");

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[0]?.dispatchEvent(new Event("click", { bubbles: true }));
const firstClass = document.body.className;
const firstMessage = liveRegion?.textContent;

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[1]?.dispatchEvent(new Event("click", { bubbles: true }));
const secondClass = document.body.className;
const secondMessage = liveRegion?.textContent;

assert.equal(firstClass, listItems[0]?.id);
assert.equal(secondClass, listItems[1]?.id);
assert.notEqual(firstClass, secondClass);
assert.notEqual(firstMessage, secondMessage);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body>

  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

#status {
  text-align: center;
  min-height: 20px;
}
```

```js
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body class="theme-light">  
  <button
    class="dropdown-btn"
    id="theme-switcher-button"
    aria-haspopup="true"
    aria-expanded="false"
    aria-controls="theme-dropdown"
  >
    Switch Theme
  </button>

  <ul
    class="dropdown-content"
    id="theme-dropdown"
    role="menu"
    aria-labelledby="theme-switcher-button"
    hidden
  >
    <li class="dropdown-item" id="theme-light" role="menuitem">Light</li>
    <li class="dropdown-item" id="theme-dark" role="menuitem">Dark</li>
    <li class="dropdown-item" id="theme-ocean" role="menuitem">Ocean</li>
    <li class="dropdown-item" id="theme-nord" role="menuitem">Nord</li>
  </ul>
  
  <div id="status" aria-live="polite"></div>
  
  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

.dropdown-btn {
  position: relative;
  padding: 10px 20px;
  border: none;
  cursor: pointer;
}

.dropdown-content {
  position: absolute;
  background-color: #ffffff;
  min-width: 160px;
  box-shadow: 0px 4px 8px rgba(0,0,0,0.2);
  z-index: 1;
}

.dropdown-item {
  width: 100%;
  padding: 12px;
  color: #000000;
  background-color: #ffffff;
  cursor: pointer;
}

.dropdown-item:hover {
  background-color: #f1f1f1;
}

#status {
  text-align: center;
  min-height: 20px;
}

/* Themes */
.theme-light {
  background-color: #ffffff;
  color: #000000;
}

.theme-light .dropdown-btn {
  background-color: #d3b3aa;
  color: #000000;
}

.theme-dark {
  background-color: #121212;
  color: #ffffff;
}

.theme-dark .dropdown-btn {
  background-color: #4f4f4f;
  color: #ffffff;
}

.theme-ocean {
  background-color: #0077be;
  color: #e0f7fa;
}

.theme-ocean .dropdown-btn {
  background-color: #193e51;
  color: #ffffff;
}

.theme-nord {
  background-color: #2e3440;
  color: #d8dee9;
}

.theme-nord .dropdown-btn {
  background-color: #596680;
  color: #ffffff;
}
```

```js
const themes = [
  {
    name: "light",
    message: "Hello sunshine — Light theme is on!"
  },
  {
    name: "dark",
    message: "The night is yours — Dark theme is on!"
  },
  {
    name: "nord",
    message: "The frost has settled - Nord theme is on!"
  },
  {
    name: "ocean",
    message: "Blue skies and high tides — Ocean theme is on!"
  },
];

const setTheme = (theme) => {
  document.body.className = "";
  document.body.classList.add(`theme-${theme.name}`);
};

const updateStatus = (theme) => {
  const status = document.getElementById("status");
  status.textContent = theme.message;
};

document.addEventListener("DOMContentLoaded", () => {
  themes.forEach((theme) => {
    const button = document.getElementById(`theme-${theme.name}`);
    
    if (button) {
      button.addEventListener("click", () => {
        setTheme(theme);
        updateStatus(theme);
        closeMenu();
      });
    }
  });

  const toggleButton = document.getElementById("theme-switcher-button");
  const menu = document.getElementById("theme-dropdown");

  toggleButton.addEventListener("click", () => {
    const isExpanded = toggleButton.getAttribute("aria-expanded") === "true";
    if (isExpanded) {
      closeMenu();
    } else {
      openMenu();
    }
  });

  function closeMenu() {
    menu.setAttribute("hidden", "");
    toggleButton.setAttribute("aria-expanded", "false");
  }

  function openMenu() {
    menu.removeAttribute("hidden");
    toggleButton.setAttribute("aria-expanded", "true");
  }

  // Close the menu when clicking outside
  document.addEventListener("click", (e) => {
    if (!toggleButton.contains(e.target) && !menu.contains(e.target)) {
      closeMenu();
    }
  });
});
```
