---
id: 699b6dd1a2961bd38491ef2e
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

Das `profile`-Objekt hatte strenge Schlüssel und Typen. Wenn Sie versucht haben, eine Eigenschaft hinzuzufügen, die nicht definiert war, oder einen inkompatiblen Typ zuzuweisen, hätte TypeScript einen Fehler geworfen.

Es wäre schön, wenn das `userRoles`-Objekt nicht so streng wäre und flexiblere Schlüssel hätte.

Wie Sie sich aus den vorherigen Lektionen erinnern, können Sie dieses Ergebnis mit dem `Record`-Typ erreichen:

```ts
const studentGrades: Record<string, number> = {
  Math: 95,
  English: 88,
  Science: 92
};
```

In diesem `studentGrades`-Beispiel muss TypeScript nicht alle Fächer im Voraus kennen. So können Sie später weitere Fächer wie `Art` oder `Music` hinzufügen. TypeScript ist nur wichtig, dass die Werte für jeden Schlüssel eine Zahl sind.

Aktualisieren Sie Ihr `userRoles`-Objekt, indem Sie einen inline `Record`-Typ hinzufügen. Setzen Sie ihn auf `<string, string>`, sodass jeder Schlüssel ein beliebiger String sein kann und jeder Wert ein String sein muss.

# --hints--

Ihr `userRoles`-Objekt sollte einen inline `Record`-Typ haben, der auf `Record<string, string>` gesetzt ist.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.annotation.matches("Record<string, string>"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles = {
--fcc-editable-region--
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

console.log(userRoles);
```
