---
id: 699b4e284291adcfcb90df47
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop üben Sie den Umgang mit primitiven Typen sowie mit Objekt- und Array-Typen, indem Sie ein typsicheres Benutzerprofil erstellen.

Beginnen Sie damit, eine Variable namens `profile` zu erstellen und ihr ein Objekt zuzuweisen. Ihr Objekt sollte diese drei Eigenschaften und Werte haben:

```md
username: "codeLearner",
age: 25,
isLoggedIn: false,
```

# --hints--

Sie sollten eine Variable namens `profile` haben.

```js
assert.exists(profile);
```

Ihre `profile`-Variable sollte ein Objekt sein.

```js
assert.isObject(profile);
```

Sie sollten eine `username`-Eigenschaft in Ihrem `profile`-Objekt haben.

```js
assert.property(profile, "username");
```

Ihre `username`-Eigenschaft sollte den Wert `"codeLearner"` haben.

```js
assert.propertyVal(profile, "username", "codeLearner");
```

Ihr `profile`-Objekt sollte eine `age`-Eigenschaft haben.

```js
assert.property(profile, "age");
```

Ihre `age`-Eigenschaft sollte den Wert `25` haben.

```js
assert.propertyVal(profile, "age", 25);
```

Ihr `profile`-Objekt sollte eine `isLoggedIn`-Eigenschaft haben.

```js
assert.property(profile, "isLoggedIn");
```

Ihre `isLoggedIn`-Eigenschaft sollte den Wert `false` haben.

```js
assert.propertyVal(profile, "isLoggedIn", false);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```
