---
id: 699b64096301fd75e3efeb1b
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Wenn Sie die Konsole erneut öffnen, sehen Sie eine andere Fehlermeldung:

```md
Property 'bio' is missing in type '{ username: string; age: number; isLoggedIn: false; }' but required in type '{ username: string; age: number; isLoggedIn: boolean; bio: string; }'.
```

Wie Sie sich aus vorherigen Lektionen erinnern, können Sie Eigenschaften als optional markieren, indem Sie das `?`-Symbol neben dem Eigenschaftsnamen verwenden, so wie hier:

```ts
const musician: {
  name: string;
  instrument: string;
  isActive: boolean;
  numberOfAlbums?: number;
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

In diesem Beispiel ist `numberOfAlbums` als optional markiert, sodass TypeScript in diesem Fall keinen Fehler wirft.

Aktualisieren Sie Ihre `bio`-Eigenschaft, sodass sie als optional markiert ist.

# --hints--

Ihre `bio`-Eigenschaft sollte als optional markiert sein.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string", isOptional: true };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
--fcc-editable-region--
  bio: string
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
