---
id: 699b70b10ae344c802d106b0
title: Schritt 12
challengeType: 1
dashedName: step-12
---

# --description--

Jetzt ist es Zeit, dem `userRoles`-Objekt noch ein paar weitere Eigenschaften hinzuzufügen.

Fügen Sie eine `moderator`-Eigenschaft mit dem Wert `"medium-access"` hinzu.

Dann fügen Sie darunter eine `guest`-Eigenschaft mit dem Wert `3` hinzu.

# --hints--

Ihr `userRoles`-Objekt sollte eine `moderator`-Eigenschaft haben.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.moderator);
```

Ihre `moderator`-Eigenschaft sollte den Wert `"medium-access"` haben.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.moderator.value.matches("'medium-access'"));
```

Ihr `userRoles`-Objekt sollte eine `guest`-Eigenschaft haben.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.guest);
```

Ihre `guest`-Eigenschaft sollte den Wert der Zahl `3` haben.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.guest.value.matches("3"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"

};
--fcc-editable-region--

console.log(userRoles);
```
