---
id: 699b71d98c820ac7ab9eea2a
title: Schritt 13
challengeType: 1
dashedName: step-13
---

# --description--

Wenn Sie die Konsole erneut öffnen, sehen Sie eine neue Fehlermeldung:

```md
Type 'number' is not assignable to type 'string'.
```

Das passiert, weil Sie versuchen, eine `guest`-Eigenschaft mit einem Zahlenwert hinzuzufügen.

Um den TypeScript-Fehler zu entfernen, ändern Sie den Wert der `guest`-Eigenschaft in `"read-only"`.

Mit dieser letzten Änderung ist Ihr Benutzerprofil-Workshop abgeschlossen!

# --hints--

Ihre `guest`-Eigenschaft für das `userRoles`-Objekt sollte den Wert `"read-only"` haben.

```js
assert.propertyVal(userRoles, "guest", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
--fcc-editable-region--
  guest: 3
--fcc-editable-region--
};

console.log(userRoles);
```

# --solutions--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
  guest: "read-only"
};

console.log(userRoles);
```
