---
id: 699b61ca8c9cb9bc1968364e
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Wenn Sie die Konsole öffnen, sollten Sie die folgende Fehlermeldung sehen:

```md
Object literal may only specify known properties, and 'mood' does not exist in type '{ username: string; age: number; isLoggedIn: boolean; }'.
```

In reinem JavaScript durften Sie beliebige Eigenschaften und Werte zum `profile`-Objekt hinzufügen. In TypeScript überprüft der Compiler jedoch, dass das Objekt nur die Eigenschaften enthält, die Sie definiert haben (in diesem Fall `username`, `age` und `isLoggedIn`). Da `mood` nicht im Typ enthalten war, behandelt TypeScript es als ungültige Eigenschaft und wirft einen Fehler.

Entfernen Sie `mood: null` aus Ihrem Code.

# --hints--

`mood: null` sollte nicht mehr in Ihrem Code vorhanden sein.

```js
assert.isUndefined(profile.mood);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
--fcc-editable-region--
  mood: null
--fcc-editable-region--
};

console.log(profile);
```
