---
id: 699b62c99ec8c63470313365
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Jetzt ist es Zeit, Ihrem `profile`-Objekt eine weitere Eigenschaft hinzuzufügen.

Beginnen Sie damit, den inline-Objekttyp zu aktualisieren, um eine `bio`-Eigenschaft mit dem Typ `string` einzuschließen.

# --hints--

Sie sollten den inline-Objekttyp aktualisieren, um eine `bio`-Eigenschaft mit dem Wert `string` einzuschließen.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
--fcc-editable-region--
  username: string;
  age: number;
  isLoggedIn: boolean;
  
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
