---
id: 699b5370a43e299b4673bcaa
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Es wäre hilfreich, die Eigenschaften des `profile`-Objekts in der Konsole zu sehen, während Sie es erstellen.

Fügen Sie eine `console.log` hinzu, die `profile` in der Konsole ausgibt.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten `profile` in der Konsole ausgeben.

```js
assert.equal(getLogs()[0], profile);
```

# --seed--

## --seed-contents--

```ts
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
}

--fcc-editable-region--

--fcc-editable-region--
```
