---
id: 699b5b7cb68b6d35afca61b7
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Obwohl es möglich ist, weiterhin Eigenschaften wie diese zum `profile`-Objekt hinzuzufügen, wäre es sinnvoll, die erlaubten Eigenschaftstypen einzuschränken. Es wäre besser, die Struktur des `profile`-Objekts im Voraus zu definieren.

In vorherigen Lektionen haben Sie gelernt, wie man mit Objekttypen wie diesem arbeitet:

```ts
const composer: { 
  name: string; 
  instrument: string; 
  isActive: boolean; 
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

In diesem Beispiel ist `composer` ein Objekt mit einem inline Objekttyp. Dieses Objekt hat drei erforderliche Eigenschaften: `name`, `instrument` und `isActive`. Die ersten beiden Eigenschaften müssen vom Typ `string` sein, und `isActive` muss vom Typ `boolean` sein.

Aktualisieren Sie Ihr bestehendes `profile`-Objekt, um einen inline Objekttyp einzuschließen. `username` sollte als `string` typisiert sein, `age` als `number` und `isLoggedIn` als `boolean`.

# --hints--

Ihr `profile`-Objekt sollte einen inline Objekttyp mit einer `username`-Eigenschaft haben, die auf `string` gesetzt ist.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "username", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

Ihr `profile`-Objekt sollte einen inline Objekttyp mit einer `age`-Eigenschaft haben, die auf `number` gesetzt ist.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "age", type: "number" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

Ihr `profile`-Objekt sollte einen inline Objekttyp mit einer `isLoggedIn`-Eigenschaft haben, die auf `boolean` gesetzt ist.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "isLoggedIn", type: "boolean" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile = {
--fcc-editable-region--
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  mood: null
}

console.log(profile);
```
