---
id: 699b65caf17a8e8d340af27d
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Jetzt ist es Zeit, die letzte Eigenschaft zum `profile` Objekt hinzuzufügen.

In vorherigen Lektionen haben Sie gelernt, wie man mit Array-Typen arbeitet. Sie haben verschiedene Array-Typen kennengelernt, wie `string[]` oder `number[]`:

```ts
const scores: number[] = [95, 87, 100];
```

In diesem Beispiel akzeptiert das `scores` Array nur Zahlen. Wenn Sie versuchen, einen anderen Typ wie `string` oder `null` hinzuzufügen, wird TypeScript einen Fehler werfen.

Aktualisieren Sie Ihr `profile` Objekt, um eine `programmingLanguages` Eigenschaft hinzuzufügen, die als `string[]` typisiert ist. Setzen Sie innerhalb des Objekt-Literals diese Eigenschaft auf `["JavaScript", "Python", "C++"]`.

# --hints--

Sie sollten eine `programmingLanguages` Eigenschaft innerhalb Ihres `profile` Objekts haben.

```js
assert.property(profile, "programmingLanguages");
```

Ihre `programmingLanguages` Eigenschaft sollte den Wert `["JavaScript", "Python", "C++"]` haben.

```js
assert.deepEqual(profile.programmingLanguages, ["JavaScript", "Python", "C++"]);
```

Ihre `programmingLanguages` Eigenschaft sollte als `string[]` typisiert sein.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "programmingLanguages", type: "string[]" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
--fcc-editable-region--
  
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
  
--fcc-editable-region--
};

console.log(profile);
```
