---
id: 699b6d91cde15459d3619c4a
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Genau wie das `profile`-Objekt wäre es schön, das `userRoles`-Objekt in der Konsole zu sehen.

Fügen Sie ein `console.log()` hinzu, um das `userRoles`-Objekt zu protokollieren.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten `userRoles` in der Konsole protokollieren.

```js
assert.equal(getLogs()[1], userRoles);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

--fcc-editable-region--

--fcc-editable-region--
```
