---
id: 699b68f94a0d154c2a7a5ce0
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

Da das `profile`-Objekt nun vollständig ist, ist es Zeit, ein neues Objekt zu erstellen, das die Benutzerrollen darstellt.

Erstellen Sie eine neue Variable namens `userRoles` und weisen Sie ihr ein Objekt zu. Ihr Objekt sollte die folgenden Eigenschaften und Werte haben:

```md
admin: "full-access",
editor: "limited-access",
viewer: "read-only"
```

# --hints--

Sie sollten eine Variable namens `userRoles` haben.

```js
assert.isDefined(userRoles);
```

Ihre `userRoles`-Variable sollte ein Objekt sein.

```js
assert.isObject(userRoles);
```

Ihr `userRoles`-Objekt sollte eine `admin`-Eigenschaft haben.

```js
assert.property(userRoles, "admin");
```

Ihre `admin`-Eigenschaft sollte den Wert `"full-access"` haben.

```js
assert.propertyVal(userRoles, "admin", "full-access");
```

Ihr `userRoles`-Objekt sollte eine `editor`-Eigenschaft haben.

```js
assert.property(userRoles, "editor");
```

Ihre `editor`-Eigenschaft sollte den Wert `"limited-access"` haben.

```js
assert.propertyVal(userRoles, "editor", "limited-access");
```

Ihr `userRoles`-Objekt sollte eine `viewer`-Eigenschaft haben.

```js
assert.property(userRoles, "viewer");
```

Ihre `viewer`-Eigenschaft sollte den Wert `"read-only"` haben.

```js
assert.propertyVal(userRoles, "viewer", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--

--fcc-editable-region--
```
