---
id: 657a0ea50da0c8d9d6d7950a
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在这个项目中，你将通过构建一个骰子游戏来学习算法思维。 总共有 6 轮，每轮玩家最多可以投掷骰子 3 次并获得分数。 

HTML 和 CSS 已为你提供。 请随意探索它们。 准备就绪后，你将需要设置 HTML 变量。 获取所有 `.die` 元素并将它们分配给 `listOfAllDice` 变量。 获取你的分数输入（`#score-options` 中的 `input` 元素）和分数跨度，并将它们分配给 `scoreInputs` 和 `scoreSpans`。 将 `#current-round` 元素分配给 `roundElement`，并将 `#current-round-rolls` 元素分配给 `rollsElement`，然后对 `#total-score` 和 `#score-history` 元素执行相同操作。 将 `#roll-dice-btn`、`#keep-score-btn`、`#rules-btn` 和 `.rules-container` 分配给具有正确格式名称的变量。

当用户点击 `Show rules` 按钮时，他们应该能够在显示和隐藏游戏规则之间切换。 创建一个变量 `isModalShowing` 来跟踪游戏规则的状态。

每次用户掷骰子时，你都需要跟踪所有骰子的值。 创建一个变量 `diceValuesArr` 来跟踪这一点。

在整个游戏过程中，你需要跟踪当前得分、总得分、掷骰子次数以及玩家所处的回合。 为此目的，声明 `rolls`、`score` 和 `round` 变量。

思考一下每个变量的起始值应该是多少。 所有这些变量都应该能够重新分配。

# --hints--

你应该将 `.die` 元素分配给 `listOfAllDice`。

```js
assert.deepEqual(listOfAllDice, document.querySelectorAll(".die"));
```

你应该将 `#score-options input` 元素分配给 `scoreInputs`。

```js
assert.deepEqual(scoreInputs, document.querySelectorAll("#score-options input"));
```

你应该将 `#score-options span` 元素分配给 `scoreSpans`。

```js
assert.deepEqual(scoreSpans, document.querySelectorAll("#score-options span"));
```

你应该将 `#current-round` 元素分配给 `roundElement`。

```js
assert.deepInclude([document.getElementById("current-round"), document.querySelector("#current-round")], roundElement);
```

你应该将 `#current-round-rolls` 元素分配给 `rollsElement`。

```js
assert.deepInclude([document.getElementById("current-round-rolls"), document.querySelector("#current-round-rolls")], rollsElement);
```

你应该将 `#total-score` 元素分配给 `totalScoreElement`。

```js
assert.deepInclude([document.getElementById("total-score"), document.querySelector("#total-score")], totalScoreElement);
```

你应该将 `#score-history` 元素分配给 `scoreHistory`。

```js
assert.deepInclude([document.getElementById("score-history"), document.querySelector("#score-history")], scoreHistory);
```

你应该将 `#roll-dice-btn` 元素分配给 `rollDiceBtn`。

```js
assert.deepInclude([document.getElementById("roll-dice-btn"), document.querySelector("#roll-dice-btn")], rollDiceBtn);
```

你应该将 `#keep-score-btn` 元素分配给 `keepScoreBtn`。

```js
assert.deepInclude([document.getElementById("keep-score-btn"), document.querySelector("#keep-score-btn")], keepScoreBtn);
```

你应该将 `#rules-btn` 元素分配给 `rulesBtn`。

```js
assert.deepInclude([document.getElementById("rules-btn"), document.querySelector("#rules-btn")], rulesBtn);
```

你应该将 `.rules-container` 元素分配给 `rulesContainer`。

```js
assert.deepEqual(rulesContainer, document.querySelector(".rules-container"));
```

你的 `isModalShowing` 变量应该具有值 `false`。

```js
assert.isFalse(isModalShowing);
```

你的 `diceValuesArr` 应该有一个空数组作为值。

```js
assert.isArray(diceValuesArr);
assert.isEmpty(diceValuesArr);
```

你的 `rolls` 变量的值应该为 `0`。

```js
assert.strictEqual(rolls, 0);
```

你的 `score` 变量的值应该为 `0`。

```js
assert.strictEqual(score, 0);
```

你的 `round` 变量应该具有值 `1`。

```js
assert.strictEqual(round, 1);
```

所有新变量都应该使用 `let` 声明。

```js
assert.match(code, /let isModalShowing/);
assert.match(code, /let diceValuesArr/);
assert.match(code, /let rolls/);
assert.match(code, /let score/);
assert.match(code, /let round/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Advanced Dice Game</title>
  <link rel="stylesheet" href="styles.css" />
</head>

<body>
  <header>
    <h1>Advanced Dice Game</h1>
    <button class="btn" id="rules-btn" type="button">Show rules</button>
    <div class="rules-container">
      <h2>Rules</h2>
      <ul>
        <li>There are total of six rounds</li>
        <li>You can only roll the dice three times per round</li>
        <li>To start the game, roll the dice</li>
        <li>
          Then, choose from one of the selected scores or roll the dice again
        </li>
        <li>
          If you choose a selected score, then you will move to the next round
        </li>
        <li>
          If you decline to choose a selected score, then you can roll the
          dice again two more times
        </li>
      </ul>
      <h2 class="points">Points</h2>
      <ul>
        <li>Three of a kind: Sum of all five dice</li>
        <li>Four of a kind: Sum of all five dice</li>
        <li>Full house: Three of a kind and a pair - 25 points</li>
        <li>
          Small straight: Four of the dice have consecutive values - 30 points
        </li>
        <li>
          Large straight: All five dice have consecutive values - 40 points
        </li>
      </ul>
    </div>
  </header>

  <main>
    <form id="game">
      <div id="dice">
        <div class="die"></div>
        <div class="die"></div>
        <div class="die"></div>
        <div class="die"></div>
        <div class="die"></div>
      </div>

      <p class="rounds-text">
        <strong>Rolls:</strong> <span id="current-round-rolls">0</span> |
        <strong>Round:</strong> <span id="current-round">1</span>
      </p>

      <div id="score-options">
        <div>
          <input type="radio" name="score-options" id="three-of-a-kind" value="three-of-a-kind" disabled />
          <label for="three-of-a-kind">Three of a kind<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="four-of-a-kind" value="four-of-a-kind" disabled />
          <label for="four-of-a-kind">Four of a kind<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="full-house" value="full-house" disabled />
          <label for="full-house">Full house<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="small-straight" value="small-straight" disabled />
          <label for="small-straight">Small straight<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="large-straight" value="large-straight" disabled />
          <label for="large-straight">Large straight<span></span></label>
        </div>

        <div>
          <input type="radio" name="score-options" id="none" value="none" disabled />
          <label for="none">None of the above<span></span></label>
        </div>
      </div>

      <button class="btn" id="keep-score-btn" type="button">
        Keep the above selected score
      </button>
      <button class="btn" id="roll-dice-btn" type="button">
        Roll the dice
      </button>
    </form>

    <div id="scores">
      <h3>Score history (Total score: <span id="total-score">0</span>)</h3>
      <ol id="score-history"></ol>
    </div>
  </main>
  <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --black: #000;
  --white: #fff;
  --grey: #3b3b4f;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

header {
  color: var(--light-grey);
  text-align: center;
}

h1 {
  font-size: 2.5rem;
  margin: 25px 0;
}

.rules-container {
  display: none;
  background-color: var(--light-grey);
  color: var(--black);
  width: 50%;
  margin: 20px auto;
  height: 300px;
  border-radius: 10px;
  overflow-y: scroll;
}

.rules-container ul {
  list-style-type: none;
}

.points {
  margin-top: 15px;
}

main {
  background-color: var(--light-grey);
  padding: 20px;
  display: grid;
  grid-template-columns: repeat(auto-fit, minmax(500px, 1fr));
  gap: 1rem;
  margin: auto;
  justify-items: center;
  width: 50%;
  border-radius: 10px;
}

#dice {
  display: flex;
  justify-content: space-around;
  margin-bottom: 15px;
}

.die {
  width: 40px;
  height: 40px;
  text-align: center;
  margin-right: 15px;
  border: 4px solid var(--black);
  padding: 10px;
}

.rounds-text {
  text-align: center;
}

input[type="radio"]:disabled + label {
  color: var(--grey);
}

#score-history {
  margin-top: 15px;
  text-align: center;
  list-style-position: inside;
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px 0 10px 0.5rem;
  color: var(--black);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover:enabled {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

@media (max-width: 992px) {
  main {
    width: 100%;
  }
}

@media (max-width: 500px) {
  .btn {
    width: 120px;
  }
}

```

```js
--fcc-editable-region--

--fcc-editable-region--
```
