---
id: acda2fb1324d9b0fa741e6b5
title: реалізуйте інструмент для перевірки закінчення рядка
challengeType: 26
dashedName: implement-a-string-ending-checker-function
---

# --description--

У цій лабораторній роботі ви реалізуєте функцію, яка перевіряє, чи закінчується рядок заданим цільовим рядком.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історія користувача:**

1. Ви повинні створити функцію з назвою `confirmEnding`, яка приймає два параметри: рядок для перевірки та рядок, з яким порівнюють.
2. Функція має повертати `true`, якщо перший рядок закінчується другим рядком, інакше — `false`.
3. Ви не повинні використовувати метод `.endsWith()`; замість цього використайте один із методів підрядка JavaScript, щоб досягти цього.

# --hints--

Ви повинні створити функцію з назвою `confirmEnding`.

```js
assert.isFunction(confirmEnding);
```

`confirmEnding` має приймати 2 параметри.

```js
assert.lengthOf(confirmEnding, 2);
```

`confirmEnding("Bastian", "n")` має повертати `true`.

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

`confirmEnding("Congratulation", "on")` має повертати `true`.

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

`confirmEnding("Connor", "n")` має повертати `false`.

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

`confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")` має повертати `false`.

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

`confirmEnding("He has to give me a new name", "name")` має повертати `true`.

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

`confirmEnding("Open sesame", "same")` має повертати `true`.

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

`confirmEnding("Open sesame", "sage")` має повертати `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

`confirmEnding("Open sesame", "game")` має повертати `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing", "mountain")` має повертати `false`.

```js
assert.isFalse(
  confirmEnding(
    "If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing",
    'mountain'
  )
);
```

`confirmEnding("Abstraction", "action")` має повертати `true`.

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

Ваш код не повинен використовувати вбудований метод `.endsWith()` для розв’язання лабораторної роботи.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
