---
id: 6606c05b5624a54ab85808fa
title: 28단계
challengeType: 20
dashedName: step-28
---

# --description--

보시다시피, 출력은 `False`입니다. 왜냐하면 `8`이 이미 보드의 첫 번째 행에 있기 때문입니다. 이제 `8`를 `7`으로 바꾸세요.

# --hints--

`print` 호출을 수정하여 `valid_in_row(0, 7)`에 `gameboard`를 호출한 결과를 출력하세요.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.valid_in_row(0, 7))")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
print(gameboard.valid_in_row(0, 8))
--fcc-editable-region--
```
