---
id: 6606bf4561f8794a0d345919
title: 27단계
challengeType: 20
dashedName: step-27
---

# --description--

이제 `valid_in_row` 메서드를 테스트할 시간입니다. `valid_in_row`에서 `gameboard`을 호출하세요. `0`과 `8`를 인수로 전달하고 결과를 출력하세요.

다시 한 번, 메서드는 세 개의 매개변수로 정의되었지만 `self`가 메서드가 호출되는 객체로 자동 전달되기 때문에 두 개의 인수만으로 호출된다는 점을 주목하세요.

# --hints--

`valid_in_row(0, 8)`에서 `gameboard`을 호출한 결과를 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.valid_in_row(0, 8))")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--

--fcc-editable-region--
```
