---
id: 660699aabc59c532f2d556e5
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

인스턴스 메서드를 호출하려면 점 표기법을 사용해야 합니다:

```py
instance_name.method_name()
```

여기서 `instance_name`는 인스턴스 또는 객체이고, `method_name`은 호출하려는 메서드입니다.

`spam` 객체의 `gameboard` 메서드를 호출하세요.

# --hints--

`spam` 객체의 `gameboard` 메서드를 `gameboard.spam()`와 함께 호출해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
