---
id: 66069b992c1c5e3451f3deb0
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

인스턴스화는 빈 객체를 만듭니다. `__init__` 메서드는 객체를 사용자 지정 상태로 인스턴스화할 수 있게 하는 특별한 메서드입니다. 클래스가 `__init__` 메서드를 구현하면, 인스턴스화 시 자동으로 `__init__`가 호출됩니다.

`Board` 클래스 안에서 `spam` 메서드를 삭제하고 `__init__` 매개변수를 포함하는 `self` 메서드로 교체하세요.

# --hints--

`spam` 클래스에 `Board` 메서드가 없어야 합니다.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

`__init__` 클래스에 `Board` 메서드를 정의해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("__init__")`)) })
```

`__init__` 메서드는 `self` 매개변수를 가져야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
