---
id: 6606a3ccb1eea93a23c066bf
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

이제 스도쿠 보드에서 빈 칸을 찾는 메서드를 작업할 차례입니다.

`Board` 클래스 안에 `find_empty_cell`이라는 빈 메서드를 만들고 `self` 매개변수를 추가하세요.

# --hints--

`find_empty_cell` 클래스 안에 `Board`라는 메서드를 정의해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("find_empty_cell")`)) })
```

`find_empty_cell` 메서드는 `self` 매개변수를 가져야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        self.board = board
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
