---
id: 6606b6b7760d0643c3b4eb29
title: 19단계
challengeType: 20
dashedName: step-19
---

# --description--

`try` 블록 안의 코드가 예외를 발생시켜도 프로그램이 계속 실행되길 원하며, `pass` 문장이 이를 수행합니다.

이 코드는 작동하지만, `except` 키워드 뒤에 예외 유형을 명시하는 것이 좋은 습관으로 여겨집니다.

`ValueError`가 발생할 수 있음을 알고 있으므로 `except` 키워드 뒤에 공백을 두고 그 뒤에 `ValueError`를 추가하세요.

# --hints--

코드에 `except ValueError:`가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_for_loops()[0].find_bodies()[0].is_equivalent("try:\\n  col = contents.index(0)\\n  return row, col\\nexcept ValueError:\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except:
                pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
