---
id: 6606c2d203a8124c83b2234b
title: 31단계
challengeType: 20
dashedName: step-31
---

# --description--

주어진 숫자가 현재 행의 지정된 열에 있는 숫자와 같지 않은지 확인해야 합니다.

이를 위해 `pass`를 `0`부터 `8`까지(포함) 범위를 반복하는 제너레이터 표현식으로 바꾸고, 각 `row`에 대해 보드의 지정된 `row`와 열 `col`에 있는 숫자가 `num`와 다른지 평가하세요.

# --hints--

`pass`를 삭제하고 제너레이터 표현식 `(self.board[row][col] != num for row in range(9))`를 만드세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_col").find_body().is_equivalent("(self.board[row][col] != num for row in range(9))")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]
--fcc-editable-region--
    def valid_in_col(self, col, num):
        pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
