---
id: 660699119472f332798860ad
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

이제 `pass`을 `print` 호출로 바꾸고 문자열 `'Spam!'`를 전달하세요.

# --hints--

`pass`을 삭제하고 `'Spam!'` 메서드 내에서 문자열 `spam`를 출력하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").find_body().is_equivalent("print('Spam!')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        pass
    
gameboard = Board()
--fcc-editable-region--
```
