---
id: 66069b0b36053733a2f012fe
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

이제 `spam` 호출을 삭제하세요.

# --hints--

코드에 `gameboard.spam()`이 없어야 합니다.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
gameboard.spam()
--fcc-editable-region--
```
