---
id: 6606cd69f56e27516583b0cc
title: 37단계
challengeType: 20
dashedName: step-37
---

# --description--

이제 보드 그리드 내 3x3 정사각형의 시작 행 인덱스를 계산하고, 각 3x3 정사각형의 시작 행 인덱스가 3의 배수인지 확인해야 합니다.

이는 정수 나눗셈 결과에 `row // 3`를 곱한 값을 취함으로써 달성할 수 있습니다. `pass`를 변수 `row_start`로 바꾸고 `(row // 3) * 3`를 할당하세요.


# --hints--

`pass`를 삭제하고 `(row // 3) * 3` 값을 가진 변수 `row_start`를 선언하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_square").find_body().is_equivalent("row_start = (row//3)*3")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))
--fcc-editable-region--
    def valid_in_square(self, row, col, num):
        pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
