---
id: 6606c3fd5634684d48a7887b
title: 32단계
challengeType: 20
dashedName: step-32
---

# --description--

이전 단계에서 작성한 제너레이터 표현식은 지정한 열의 모든 행에 대해 `self.board[row][col] != num` 조건이 `True`인지 `False`인지 나타내는 불리언 값 목록을 생성합니다.

그 제너레이터 표현식을 `all()` 함수에 전달하여 열의 모든 요소가 `num`와 다른지 확인하고 결과를 반환하세요.

# --hints--

제너레이터 표현식을 `all()` 호출의 인수로 전달하고 결과를 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_col").find_body().is_equivalent("return all(self.board[row][col] != num for row in range(9))")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]
--fcc-editable-region--
    def valid_in_col(self, col, num):
        (self.board[row][col] != num for row in range(9))
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
