---
id: 6606cc754a8834509cd0afb6
title: 36단계
challengeType: 20
dashedName: step-36
---

# --description--

또 확인해야 할 것은 숫자가 3x3 정사각형에 삽입될 수 있는지 여부입니다.

`Board` 클래스 안에 `valid_in_square`라는 이름의 메서드를 네 개의 매개변수 `self`, `row`, `col`, `num`로 만드세요. 여기서 `row`, `col`, `num`는 각각 행 인덱스, 열 인덱스, 그리고 확인할 숫자를 나타냅니다.

# --hints--

`valid_in_square` 클래스 내에 `Board`라는 새 메서드를 만드세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_square")`)) })
```

`valid_in_square` 메서드는 네 개의 매개변수 `self`, `row`, `col`, `num`를 이 순서대로 가져야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_square").has_args("self, row, col, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
