---
id: 6606d8c323d6205890fbbd54
title: 46단계
challengeType: 20
dashedName: step-46
---

# --description--

`Board` 클래스 내에 `is_valid`라는 또 다른 메서드를 만들고 세 개의 매개변수 `self`, `empty`, `num`를 지정하세요. 여기서 `empty`는 빈 칸의 행과 열 인덱스를 나타내는 튜플이고 `num`는 확인할 숫자입니다.

이 메서드는 주어진 숫자가 스도쿠 보드의 빈 칸에 대해 유효한 선택인지, 즉 지정된 빈 칸의 행, 열, 그리고 3x3 정사각형과 호환되는지 검증합니다.

# --hints--

`is_valid` 클래스 내에 `Board`라는 새 메서드를 만들어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("is_valid")`)) })
```

`is_valid` 메서드는 세 개의 매개변수 `self`, `empty`, `num`를 이 순서대로 가져야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("is_valid").has_args("self, empty, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))

    def valid_in_square(self, row, col, num):
        row_start = (row // 3) * 3
        col_start = (col // 3) * 3
        for row_no in range(row_start, row_start + 3):
            for col_no in range(col_start, col_start + 3):
                if self.board[row_no][col_no] == num:
                    return False
        return True
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
