---
id: 66069167b3307b2f4067b22b
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

클래스 이름 뒤에 소괄호 한 쌍을 붙이는 함수 표기법을 사용하여 클래스의 새 인스턴스를 만듭니다.

클래스 정의 외부에서 `Board` 클래스의 인스턴스를 만들고 `gameboard`이라는 변수에 할당하세요.

# --hints--

`gameboard` 클래스 외부에 `Board` 변수를 선언해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("gameboard")`)) })
```

`gameboard` 변수는 `Board()` 값을 가져야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
--fcc-editable-region--
```
