---
id: 6606beade9200b49aaeecd94
title: 26단계
challengeType: 20
dashedName: step-26
---

# --description--

만약 `num`가 행에 없다면, 표현식은 `True`로 평가되고 이는 숫자가 삽입하기에 유효하다는 뜻입니다.

만약 `num`가 행에 있다면, 표현식은 `False`로 평가되고 삽입은 규칙을 위반하게 됩니다.

숫자의 유효성을 확인할 수 있도록 `return` 메서드 본문 안의 표현식 앞에 `valid_in_row` 키워드를 붙이세요.

# --hints--

`num not in self.board[row]`를 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_row").find_body().is_equivalent("return num not in self.board[row]")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None
--fcc-editable-region--
    def valid_in_row(self, row, num):
        num not in self.board[row]
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
