---
id: 66069e5759b800364707988e
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

`__init__` 메서드로 돌아가면, 해결할 퍼즐을 나타내는 추가 매개변수가 필요합니다.

`board` 메서드에 `__init__`이라는 두 번째 매개변수를 추가하고, 함수 호출 시 인수를 전달하듯이 `gameboard` 리스트를 전달하여 `puzzle`의 인스턴스화를 수정하세요.

# --hints--

`__init__` 메서드는 `self`, `board` 순서로 두 개의 매개변수를 가져야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self, board")`)) })
```

`gameboard` 변수는 `Board(puzzle)` 값을 가져야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board(puzzle)")`)) })
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board()
--fcc-editable-region--
```
