---
id: 6606c14182435d4bab0de2ee
title: 30단계
challengeType: 20
dashedName: step-30
---

# --description--

다음으로, 숫자가 해당 열에 이미 존재하지 않는지 확인하여 스도쿠 보드의 지정된 열에 숫자를 삽입할 수 있는지 검사하는 메서드를 만드세요.

`Board` 클래스 내에 `valid_in_col`라는 이름의 메서드를 만들고 세 개의 매개변수 `self`, `col`, `num`를 지정하세요.
여기서 `col`와 `num`는 각각 열 인덱스와 검사할 숫자입니다.

# --hints--

`valid_in_col` 클래스 내에 `Board`라는 새 메서드를 만드세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_col")`)) })
```

`valid_in_col` 메서드는 세 개의 매개변수 `self`, `col`, `num`를 이 순서대로 가져야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_col").has_args("self, col, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
