---
id: 6606a219f9efbf38ad496f67
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

점 표기법을 사용해 인스턴스 속성에 접근할 수도 있습니다.

`Board` 클래스 외부에서 `gameboard` 객체를 초기화한 후, `gameboard.board`를 사용해 `board` 객체의 `gameboard` 속성에 접근하고 결과를 화면에 출력하세요.

# --hints--

`gameboard.board`를 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.board)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        self.board = board

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
