---
id: 66069f86f58f85371d47123e
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

속성은 객체와 연관된 변수로, 일반 변수처럼 데이터를 저장하는 데 사용합니다.

`__init__` 메서드 안에서 `board` 클래스의 인스턴스를 생성할 때 전달되는 `Board` 매개변수를 `board`를 사용해 인스턴스 속성 `self.board`에 할당하세요.

`self.board`는 클래스 인스턴스의 `board` 속성을 가리킵니다. 이는 `Board` 클래스에서 생성된 객체에 속한 변수입니다.

# --hints--

`pass`를 삭제하고 `board` 메서드 안에서 `self.board` 매개변수를 `__init__`에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").find_body().is_equivalent("self.board = board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
