---
id: 6606bc4e5535c0484990ccd5
title: 24단계
challengeType: 20
dashedName: step-24
---

# --description--

다음으로, 주어진 숫자가 스도쿠 보드의 특정 행에 삽입될 수 있는지 확인하는 메서드를 작업할 거예요.

`Board` 클래스 내에 `valid_in_row`라는 이름의 메서드를 만들고 세 개의 매개변수 `self`, `row`, `num`를 주세요. 여기서 `self`는 클래스 인스턴스를 나타내고, `row`와 `num`는 각각 행 인덱스와 확인할 숫자입니다.


# --hints--

`valid_in_row` 클래스 내에 `Board`라는 새 메서드를 만들어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_row")`)) })
```

`valid_in_row` 메서드는 세 개의 매개변수 `self`, `row`, `num`를 이 순서대로 가져야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_row").has_args("self, row, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
