---
id: 6606b961ebcf04460f8af76e
title: 21단계
challengeType: 20
dashedName: step-21
---

# --description--

`find_empty_cell` 메서드가 제대로 작동하는지 `gameboard`에서 호출하고 결과를 출력해서 테스트하세요.

`find_empty_cell`가 매개변수 하나로 정의되어 있지만, 메서드를 호출하는 객체가 자동으로 `self`에 전달되므로 함수 호출 시 인수를 전달해 값을 주면 안 됩니다.

# --hints--

`gameboard.find_empty_cell()`를 출력하세요.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.find_empty_cell())")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--

--fcc-editable-region--
```
