---
id: 66068fb0bfddba2b7977eb60
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 프로젝트에서는 스도쿠 퍼즐 해결기를 만들어 클래스를 배우고 객체를 이해합니다.

Python에서 클래스는 객체를 만들기 위한 청사진입니다. 클래스에서 만들어진 객체는 해당 클래스의 인스턴스입니다. 다음 구문으로 클래스를 만들 수 있습니다:

```py
class ClassName:
    pass
```

여기서 `class`는 클래스를 정의할 때 필요한 키워드이고 `ClassName`는 관례에 따라 *PascalCase*로 작성된 클래스 이름입니다.

`Board` 클래스를 만들어 시작하세요.

# --hints--

`Board`라는 이름의 클래스를 만들어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
