---
id: 6606b224a69a293f98f8db8f
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

`.index()` 메서드는 값이 발견되지 않을 때 `ValueError` 예외를 발생시킵니다. 프로그램이 실행을 멈추지 않도록 하려면 이 코드 줄을 `try` 블록 안에 중첩하세요. `try` 문은 예외가 발생할 수 있는 코드를 감싸는 데 사용됩니다. 반면 `except` 절은 예외가 발생했을 때 실행할 대체 코드를 제공합니다:

```py
try:
    <code>
except:
    <code>
```

`col` 할당을 `try` 블록 안에 넣으세요. 그런 다음 `except` 절을 만들고 그 본문을 `pass`로 채우세요.

# --hints--

`col` 할당을 `try` 블록 안에 넣고 `except` 절을 만드세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_for_loops()[0].find_bodies()[0].is_equivalent("try:\\n  col = contents.index(0)\\nexcept:\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            col = contents.index(0)
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
