---
id: 6606927d010be4300a4e5330
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

인스턴스화는 빈 객체를 만듭니다. 하지만 클래스는 각 인스턴스에 대한 로컬 함수 같은 메서드를 가질 수 있습니다. 클래스 내에서 메서드는 다음과 같이 선언합니다:

```py
class ClassName:
    def method_name():
        pass
```

`Board` 클래스 안에서 `pass`을 빈 메서드 `spam`로 바꾸세요.

# --hints--

`spam` 클래스 안에 `Board`라는 이름의 메서드를 정의해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
    
gameboard = Board()
--fcc-editable-region--
```
