---
id: 6606933d6813a8308c962dd1
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

인스턴스 메서드가 되려면, 메서드는 관례상 `self`라는 특별한 매개변수를 필요로 합니다. 이 매개변수는 클래스 인스턴스에 대한 참조이며 항상 첫 번째 매개변수여야 합니다.

`self` 메서드에 `spam` 매개변수를 추가하세요.

# --hints--

`spam` 메서드에는 `self` 매개변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam():
        pass
    
gameboard = Board()
--fcc-editable-region--
```
