---
id: 6606a4641ec48b3a9fe8c2fc
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

`enumerate` 내장 함수는 반복 가능한 객체를 인수로 받아 순회할 수 있는 enumerate 객체를 반환합니다.
이 함수는 기본적으로 0부터 시작하는 개수와 반복 가능한 객체의 값을 제공합니다.

```py
iterable = ['a', 'b', 'c']
for i, j in enumerate(iterable):
    print(i, j)
```

위 예제의 루프는 튜플 `0, a`, `1, b`, 그리고 `2, c`를 출력할 것입니다.

`find_empty_cell` 메서드 안에서 `pass`를 `for` 루프로 바꾸세요. 이 루프는 `enumerate()` 함수를 사용해 스도쿠 보드의 각 행을 순회합니다. 현재 행의 인덱스는 `row`로, 현재 행의 요소들은 `contents`로 사용하세요.


# --hints--

`pass`를 `for` 루프로 바꾸어 `enumerate(self.board)`를 순회해야 합니다. 루프 변수로 `row`과 `contents`를 사용하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_body().is_equivalent("for row, contents in enumerate(self.board):\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
