---
id: 660a737f0f72b51de361051c
title: 54단계
challengeType: 20
dashedName: step-54
---

# --description--

`:=` 연산자는 표현식 중간에 변수를 할당할 수 있는 기능을 제공합니다. 구문은 `name := val`이며, 여기서 `name`는 변수 이름이고 `val`는 변수 값입니다.

이 구조는 공식적으로 *할당 표현식*이라고 하며, `:=` 연산자는 일반적으로 *월러스* 연산자라고 부릅니다.

`self.find_empty_cell()` 호출을 여러 번 사용할 것이므로, 월러스 연산자를 사용해 `next_empty` 변수에 할당하세요. 그런 다음, 할당을 소괄호 한 쌍으로 감싸세요.

이렇게 하면 할당과 조건 검사를 한 줄로 결합하여 코드를 더 간결하게 만들 수 있습니다.

# --hints--

`if` 조건을 `(next_empty := self.find_empty_cell()) is None`로 수정하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("solver").find_ifs()[0].find_conditions()[0].is_equivalent("(next_empty := self.find_empty_cell()) is None")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))

    def valid_in_square(self, row, col, num):
        row_start = (row // 3) * 3
        col_start = (col // 3) * 3
        for row_no in range(row_start, row_start + 3):
            for col_no in range(col_start, col_start + 3):
                if self.board[row_no][col_no] == num:
                    return False
        return True

    def is_valid(self, empty, num):
        row, col = empty
        valid_in_row = self.valid_in_row(row, num)
        valid_in_col = self.valid_in_col(col, num)
        valid_in_square = self.valid_in_square(row, col, num)
        return all([valid_in_row, valid_in_col, valid_in_square])
--fcc-editable-region--
    def solver(self):
        if self.find_empty_cell() is None:
            return True
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
