---
id: 657efa642593c5746acc5c81
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Перетворіть будь-який символ у верхньому регістрі в тілі інструкції `if` на нижній регістр та додайте перед ним знак підкреслення.

Використайте метод `.lower()`, щоб перетворити верхній регістр символа на нижній. Потім додайте знак підкреслення перед символом. Призначте результати до змінної під назвою `converted_character`.

# --hints--

You should assign the modified character to a variable named `converted_character`.

```js
({
    test: () => {
   assert(runPython(`
      converted_character_variable = (
          _Node(_code)
          .find_function("convert_to_snake_case")
          .find_for_loops()[0]
          .find_ifs()[0]
          .find_variable("converted_character")
      )

      converted_character_variable.is_equivalent(
          "converted_character = '_' + char.lower()"
      ) or converted_character_variable.is_equivalent(
          "converted_character = f'_{char.lower()}'"
      )
    `)); 
    }
})
```

You should not have `pass` in your `if` statement.

```js
({
    test: () => {
        assert(runPython(`not _Node(_code).find_function("convert_to_snake_case").find_for_loops()[0].find_ifs()[0].has_stmt("pass")`)); 
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            pass
--fcc-editable-region--
```
