---
id: afcc8d540bea9ea2669306b6
title: 文字列を繰り返す関数を作成する
challengeType: 26
dashedName: build-a-string-repeating-function
---

# --description--

このラボでは、指定された文字列を特定の回数だけ繰り返す関数を作成します。このラボの目的のために、組み込みの`.repeat()`メソッドは使用しないでください。

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー:**

1. `repeatStringNumTimes`という名前の関数を作成し、2つのパラメータ（文字列と数値）を受け取る必要があります。
2. 関数は指定された回数だけ文字列を繰り返して返す必要があります。
3. 数値が0以下の場合、関数は空の文字列を返す必要があります。

# --hints--

`repeatStringNumTimes`という名前の関数を作成してください。

```js
assert.isFunction(repeatStringNumTimes);
```

`repeatStringNumTimes`は2つのパラメータを受け取る必要があります。

```js
assert.lengthOf(repeatStringNumTimes, 2);
```

関数`repeatStringNumTimes`は常に文字列を返す必要があります。

```js
assert.isString(repeatStringNumTimes('hello', 3));
```

`repeatStringNumTimes("*", 3)`は文字列`***`を返す必要があります。

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

`repeatStringNumTimes("abc", 3)`は文字列`abcabcabc`を返す必要があります。

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

`repeatStringNumTimes("abc", 4)`は文字列`abcabcabcabc`を返す必要があります。

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

`repeatStringNumTimes("abc", 1)`は文字列`abc`を返す必要があります。

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

`repeatStringNumTimes("*", 8)`は文字列`********`を返す必要があります。

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

`repeatStringNumTimes("abc", -2)`は空の文字列（`""`）を返す必要があります。

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

`repeatStringNumTimes("abc", 0)`は`""`を返す必要があります。

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

組み込みの`repeat()`メソッドは使用しないでください。

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
