---
id: bad87fee1348bd9aedf08736
title: стилізуйте елемент body у HTML
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB77PHW'
forumTopicId: 18313
dashedName: style-the-html-body-element
---

# --description--

Тепер почнемо з чистого аркуша і поговоримо про успадкування CSS.

Кожна HTML-сторінка має елемент `body`.

# --instructions--

Ми можемо довести, що елемент `body` існує тут, задавши йому проп `background-color` чорного кольору.

Ми можемо зробити це, додавши наступне до нашого елемента `style`:

```css
body {
  background-color: black;
}
```

# --hints--

Ваш елемент `body` повинен мати проп `background-color` чорного кольору.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color'];

assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

Ваш CSS-правило має бути правильно відформатоване з відкриваючими та закриваючими фігурними дужками.

```js
assert.match(code, /<style>\s*body\s*\{\s*background.*\s*:\s*.*;\s*\}\s*<\/style>/i);
```

Ваш CSS-правило має закінчуватися крапкою з комою.

```js
assert.match(code, /<style>\s*body\s*\{\s*background.*\s*:\s*.*;\s*\}\s*<\/style>/i);
```

# --seed--

## --seed-contents--

```html
<style>

</style>
```

# --solutions--

```html
<style>
body {
  background-color: black;
}
</style>
```
