---
id: bad87fee1348bd9aede08718
title: Використання значень RGB для фарбування елементів
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkp2fr'
forumTopicId: 18369
dashedName: use-rgb-values-to-color-elements
---

# --description--

Ще один спосіб представити кольори в CSS — це використання значень `RGB`.

Значення `RGB` для чорного виглядає так:

```css
rgb(0, 0, 0)
```

Значення `RGB` для білого виглядає так:

```css
rgb(255, 255, 255)
```

Замість того, щоб використовувати шість шістнадцяткових цифр, як у hex-коді, у `RGB` ви вказуєте яскравість кожного кольору числом від 0 до 255.

Якщо порахувати, дві цифри для одного кольору дорівнюють 16 помножити на 16, що дає нам 256 можливих значень. Тож `RGB`, який рахує від нуля, має точно таку ж кількість можливих значень, як і hex-код.

Ось приклад, як змінити фон `body` на помаранчевий, використовуючи його RGB-код.

```css
body {
  background-color: rgb(255, 165, 0);
}
```

# --instructions--

Давайте замінимо hex-код у кольорі фону нашого елемента `body` на RGB-значення чорного: `rgb(0, 0, 0)`

# --hints--

Ваш елемент `body` має мати чорний фон.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

Ви маєте використати `rgb`, щоб задати чорний фон елементу `body`.

```js
assert.match(__helpers.removeCssComments(code), /rgb\s*\(\s*0\s*,\s*0\s*,\s*0\s*\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #F00;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: rgb(0, 0, 0);
  }
</style>
```
