---
id: bad87fee1348bd9aedf08721
title: Використання шістнадцяткового коду для змішування кольорів
challengeType: 0
videoUrl: 'https://scrimba.com/c/cK89PhP'
forumTopicId: 18359
dashedName: use-hex-code-to-mix-colors
---

# --description--

Нагадаємо, що шістнадцяткові коди використовують 6 шістнадцяткових знаків для представлення кольорів, по два для червоного (R), зеленого (G) та синього (B) компонентів.

З цих трьох чистих кольорів (червоний, зелений і синій) ми можемо варіювати кількість кожного, щоб створити понад 16 мільйонів інших кольорів!

Наприклад, помаранчевий — це чистий червоний, змішаний з деякою кількістю зеленого, і без синього. У шістнадцятковому коді це відповідає `#FFA500`.

Цифра `0` — найнижче число в шістнадцятковому коді, і вона означає повну відсутність кольору.

Цифра `F` — найвище число в шістнадцятковому коді, і вона означає максимальну можливу яскравість.

# --instructions--

Замініть слова кольорів у нашому `style` елементі на їхні правильні шістнадцяткові коди.

<table><tbody><tr><th>Колір</th><th>Шістнадцятковий код</th></tr><tr><td>Dodger Blue</td><td><code>#1E90FF</code></td></tr><tr><td>Зелений</td><td><code>#00FF00</code></td></tr><tr><td>Помаранчевий</td><td><code>#FFA500</code></td></tr><tr><td>Червоний</td><td><code>#FF0000</code></td></tr></tbody></table>

# --hints--

Вашому `h1` елементу з текстом `I am red!` слід надати `color` червоний.

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Для кольору червоний слід використати `hex code` замість слова `red`.

```js
assert.match(code, /\.red-text\s*?{\s*?color\s*:\s*?(#FF0000|#F00)\s*?;?\s*?}/gi);
```

Вашому `h1` елементу з текстом `I am green!` слід надати `color` зелений.

```js
const greenText = document.querySelector('.green-text');
const color = window.getComputedStyle(greenText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 0)');
```

Для кольору зелений слід використати `hex code` замість слова `green`.

```js
assert.match(code, /\.green-text\s*?{\s*?color\s*:\s*?(#00FF00|#0F0)\s*?;?\s*?}/gi);
```

Вашому `h1` елементу з текстом `I am dodger blue!` слід надати `color` dodger blue.

```js
const blueText = document.querySelector('.dodger-blue-text');
const color = window.getComputedStyle(blueText)['color']; 
assert.strictEqual(color, 'rgb(30, 144, 255)');
```

Для кольору dodger blue слід використати `hex code` замість слова `dodgerblue`.

```js
assert.match(code, /\.dodger-blue-text\s*?{\s*?color\s*:\s*?#1E90FF\s*?;?\s*?}/gi);
```

Вашому `h1` елементу з текстом `I am orange!` слід надати `color` помаранчевий.

```js
const orangeText = document.querySelector('.orange-text');
const color = window.getComputedStyle(orangeText)['color']; 

assert.strictEqual(color, 'rgb(255, 165, 0)');
```

Для кольору помаранчевий слід використати `hex code` замість слова `orange`.

```js
assert.match(code, /\.orange-text\s*?{\s*?color\s*:\s*?#FFA500\s*?;?\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: black;
  }
  .green-text {
    color: black;
  }
  .dodger-blue-text {
    color: black;
  }
  .orange-text {
    color: black;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="green-text">I am green!</h1>

<h1 class="dodger-blue-text">I am dodger blue!</h1>

<h1 class="orange-text">I am orange!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: #FF0000;
  }
  .green-text {
    color: #00FF00;
  }
  .dodger-blue-text {
    color: #1E90FF;
  }
  .orange-text {
    color: #FFA500;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="green-text">I am green!</h1>

<h1 class="dodger-blue-text">I am dodger blue!</h1>

<h1 class="orange-text">I am orange!</h1>
```
