---
id: bad88fee1348bd9aedf08825
title: Відрегулювати відступ елемента
challengeType: 0
videoUrl: 'https://scrimba.com/c/cED8ZC2'
forumTopicId: 301083
dashedName: adjust-the-padding-of-an-element
---


# --description--

Тепер давайте відкладемо наш застосунок Cat Photo на деякий час і дізнаємося більше про стилізацію HTML.

Ви, можливо, вже помітили, що всі HTML-елементи по суті є маленькими прямокутниками.

Три важливі властивості контролюють простір, який оточує кожен HTML-елемент: `padding`, `border` та `margin`.

`padding` елемента контролює кількість простору між вмістом елемента та його `border`.

Тут ми бачимо, що синій блок і червоний блок вкладені у жовтий блок. Зверніть увагу, що червоний блок має більше `padding`, ніж синій блок.

Коли ви збільшуєте `padding` синього блоку, це збільшує відстань (`padding`) між текстом і кордоном навколо нього.

# --instructions--

Змініть `padding` вашого синього блоку, щоб він відповідав `padding` вашого червоного блоку.

# --hints--

Ваш клас `blue-box` повинен надавати елементам `20px` `padding`.

```js
const blueBox = document.querySelector('.blue-box');
const paddingTop = window.getComputedStyle(blueBox)['padding-top'];
assert.strictEqual(paddingTop, '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 10px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
