---
id: bad87fee1348bd9aedf08746
title: Наслідування стилів від елемента body
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bmdtR'
forumTopicId: 18204
dashedName: inherit-styles-from-the-body-element
---

# --description--

Тепер ми довели, що кожна HTML-сторінка має елемент `body`, і що його елемент `body` також можна стилізувати за допомогою CSS.

Пам’ятайте, ви можете стилізувати свій елемент `body` так само, як і будь-який інший HTML-елемент, і всі ваші інші елементи успадкують стилі вашого елемента `body`.

# --instructions--

Спочатку створіть елемент `h1` з текстом `Hello World`.

Потім надайте всім елементам на вашій сторінці колір `green`, додавши `color: green;` до оголошення стилю вашого елемента `body`.

Нарешті, надайте вашому елементу `body` шрифт `monospace`, додавши `font-family: monospace;` до оголошення стилю вашого елемента `body`.

# --hints--

Ви повинні створити елемент `h1`.

```js
assert.isNotEmpty(document.querySelectorAll('h1'));
```

Ваш елемент `h1` має містити текст `Hello World`.

```js
assert.match(
  document.querySelector('h1').textContent,
  /hello world/i
);
```

Ваш елемент `h1` має мати кінцевий тег.

```js
const commentlessCode = __helpers.removeHtmlComments(code);
assert.match(commentlessCode, /<\/h1>/g);
assert.match(commentlessCode, /<h1/g);
assert.lengthOf(commentlessCode.match(/<\/h1>/g), commentlessCode.match(/<h1/g).length);
```

Ваш елемент `body` має мати властивість `color` зі значенням `green`.

```js
const bodyElement = document.querySelector('body');
const color = window.getComputedStyle(bodyElement)['color']; 
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

Ваш елемент `body` має мати властивість `font-family` зі значенням `monospace`.

```js
const bodyElement = document.querySelector('body');
const fontFamily = window.getComputedStyle(bodyElement)['font-family'];
assert.include(fontFamily.toLowerCase(), "monospace");
```

Ваш елемент `h1` має успадкувати шрифт `monospace` від вашого елемента `body`.

```js
const h1Element = document.querySelector('h1');
const fontFamily = window.getComputedStyle(h1Element)['font-family'];
assert.include(fontFamily.toLowerCase(), "monospace");
```

Ваш елемент `h1` має успадкувати колір `green` від вашого елемента `body`.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }

</style>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }

</style>
<h1>Hello World!</h1>
```
