---
id: bad82fee1322bd9aedf08721
title: Розуміння абсолютних і відносних одиниць
challengeType: 0
videoUrl: 'https://scrimba.com/c/cN66JSL'
forumTopicId: 301089
dashedName: understand-absolute-versus-relative-units
---

# --description--

Останні кілька завдань встановлювали поле або відступ елемента у пікселях (`px`). Пікселі — це тип одиниці довжини, яка вказує браузеру, як розмірювати або розміщувати елемент. Окрім `px`, у CSS є кілька різних варіантів одиниць довжини, які ви можете використовувати.

Два основні типи одиниць довжини — абсолютні та відносні. Абсолютні одиниці пов’язані з фізичними одиницями довжини. Наприклад, `in` і `mm` позначають дюйми та міліметри відповідно. Абсолютні одиниці довжини приблизно відповідають фактичному виміру на екрані, але є деякі відмінності залежно від роздільності екрану.

Відносні одиниці, такі як `em` або `rem`, залежать від іншого значення довжини. Наприклад, `em` базується на розмірі шрифту елемента. Якщо ви використовуєте її для встановлення властивості `font-size` самого елемента, вона відносна до `font-size` батьківського елемента.

**Примітка:** Існує кілька варіантів відносних одиниць, які пов’язані з розміром вікна перегляду. Вони розглядаються в розділі Принципи адаптивного вебдизайну.

# --instructions--

Додайте властивість `padding` до елемента з класом `red-box` і встановіть її в `1.5em`.

# --hints--

Ваш клас `red-box` має мати властивість `padding`.

```js
const redBox =document.querySelector('.red-box'); 
const style = window.getComputedStyle(redBox); 

assert.notEqual(style['padding-top'], '0px');
assert.notEqual(style['padding-right'], '0px');
assert.notEqual(style['padding-bottom'], '0px');
assert.notEqual(style['padding-left'], '0px');
```

Ваш клас `red-box` має задавати елементам 1.5em `padding`.

```js
assert.match(__helpers.removeCssComments(code), /\.red-box\s*?{[\s\S]*padding\s*:\s*?1\.5em/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: red;
    margin: 20px 40px 20px 40px;

  }

  .green-box {
    background-color: green;
    margin: 20px 40px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box green-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: red;
    margin: 20px 40px 20px 40px;
    padding: 1.5em;
  }

  .green-box {
    background-color: green;
    margin: 20px 40px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box green-box">padding</h5>
</div>
```
