---
id: bad87fee1348bd8aedf06756
title: Заміщення оголошень класів стилізацією атрибутів id
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkpDhB'
forumTopicId: 18251
dashedName: override-class-declarations-by-styling-id-attributes
---

# --description--

Ми щойно довели, що браузери читають CSS зверху вниз у порядку їх оголошення. Це означає, що у разі конфлікту браузер використає те оголошення CSS, яке йде останнім. Зверніть увагу, що навіть якби ми поставили `blue-text` перед `pink-text` у класах елемента `h1`, браузер все одно дивився б на порядок оголошень, а не на порядок їх використання!

Але це ще не все. Існують інші способи, як можна заміщати CSS. Ви пам’ятаєте атрибути id?

Давайте замістимо ваші класи `pink-text` і `blue-text`, і зробимо елемент `h1` помаранчевим, додавши елементу `h1` id, а потім стилізувавши цей id.

# --instructions--

Додайте елементу `h1` атрибут `id` зі значенням `orange-text`. Пам’ятайте, що стилі для id виглядають так:

```html
<h1 id="orange-text">
```

Залиште класи `blue-text` і `pink-text` на вашому елементі `h1`.

Створіть CSS-оголошення для id `orange-text` у вашому елементі `style`. Ось приклад, як це виглядає:

```css
#brown-text {
  color: brown;
}
```

**Примітка:** Не має значення, чи оголосите ви цей CSS вище чи нижче класу `pink-text`, оскільки атрибут `id` завжди матиме пріоритет.

# --hints--

Ваш елемент `h1` має мати клас `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Ваш елемент `h1` має мати клас `blue-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

Ваш елемент `h1` має мати id `orange-text`.

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'),'orange-text');
```

Повинен бути лише один елемент `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Ваш id `orange-text` має мати CSS-оголошення.

```js
assert.match(__helpers.removeCssComments(code), /#orange-text\s*{/gi);
```

Ваш `h1` не повинен мати жодних атрибутів `style`.

```js
assert.notMatch(__helpers.removeHtmlComments(code), /<h1.*style.*>/gi);
```

Ваш елемент `h1` має бути помаранчевим.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 165, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 class="pink-text blue-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
  #orange-text {
    color: orange;
  }  
</style>
<h1 id="orange-text"  class="pink-text blue-text">Hello World!</h1>
```
