---
id: bad87fee1348bd9aedf08719
title: Використання скороченого шістнадцяткового коду
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkpKAm'
forumTopicId: 18338
dashedName: use-abbreviated-hex-code
---

# --description--

Багато людей відчувають себе перевантаженими через можливості понад 16 мільйонів кольорів. І складно запам’ятати шістнадцятковий код. На щастя, його можна скоротити.

Наприклад, шістнадцятковий код червоного `#FF0000` можна скоротити до `#F00`. Ця скорочена форма дає по одній цифрі для червоного, зеленого та синього.

Це зменшує загальну кількість можливих кольорів приблизно до 4 000. Але браузери інтерпретують `#FF0000` і `#F00` як точно той самий колір.

# --instructions--

Спробуйте використати скорочені шістнадцяткові коди, щоб пофарбувати правильні елементи.

<table><tbody><tr><th>Колір</th><th>Скорочений шістнадцятковий код</th></tr><tr><td>Бірюзовий</td><td><code>#0FF</code></td></tr><tr><td>Зелений</td><td><code>#0F0</code></td></tr><tr><td>Червоний</td><td><code>#F00</code></td></tr><tr><td>Фуксія</td><td><code>#F0F</code></td></tr></tbody></table>

# --hints--

Вашому `h1` елементу з текстом `I am red!` слід надати `color` червоний.

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Замість шістнадцяткового коду `#FF0000` слід використати скорочений шістнадцятковий код для кольору червоний.

```js
assert.match(code,/\.red-text\s*?{\s*?color\s*:\s*?#F00\s*?;?\s*?}/gi);
```

Вашому `h1` елементу з текстом `I am green!` слід надати `color` зелений.

```js
const greenText = document.querySelector('.green-text');
const color = window.getComputedStyle(greenText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 0)');
```

Замість шістнадцяткового коду `#00FF00` слід використати скорочений шістнадцятковий код для кольору зелений.

```js
assert.match(__helpers.removeCssComments(code), /\.green-text\s*?{\s*?color\s*:\s*?#0F0\s*?;?\s*?}/gi);
```

Вашому `h1` елементу з текстом `I am cyan!` слід надати `color` бірюзовий.

```js
const cyanText = document.querySelector('.cyan-text');
const color = window.getComputedStyle(cyanText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 255)');
```

Замість шістнадцяткового коду `#00FFFF` слід використати скорочений шістнадцятковий код для кольору бірюзовий.

```js
assert.match(__helpers.removeCssComments(code), /\.cyan-text\s*?{\s*?color\s*:\s*?#0FF\s*?;?\s*?}/gi);
```

Вашому `h1` елементу з текстом `I am fuchsia!` слід надати `color` фуксія.

```js
const fuchsiaText = document.querySelector('.fuchsia-text');
const color = window.getComputedStyle(fuchsiaText)['color']; 
assert.strictEqual(color,'rgb(255, 0, 255)');
```

Замість шістнадцяткового коду `#FF00FF` слід використати скорочений шістнадцятковий код для кольору фуксія.

```js
assert.match(__helpers.removeCssComments(code), /\.fuchsia-text\s*?{\s*?color\s*:\s*?#F0F\s*?;?\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: #000000;
  }
  .fuchsia-text {
    color: #000000;
  }
  .cyan-text {
    color: #000000;
  }
  .green-text {
    color: #000000;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="fuchsia-text">I am fuchsia!</h1>

<h1 class="cyan-text">I am cyan!</h1>

<h1 class="green-text">I am green!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: #F00;
  }
  .fuchsia-text {
    color: #F0F;
  }
  .cyan-text {
    color: #0FF;
  }
  .green-text {
    color: #0F0;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="fuchsia-text">I am fuchsia!</h1>

<h1 class="cyan-text">I am cyan!</h1>

<h1 class="green-text">I am green!</h1>
```
