---
id: bad87fee1348bd9aedf06756
title: Заміщення оголошень класів рядковими стилями
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJDRha'
forumTopicId: 18252
dashedName: override-class-declarations-with-inline-styles
---

# --description--

Отже, ми довели, що оголошення ідентифікатора заміщають оголошення класу, незалежно від того, де вони оголошені у CSS вашого `style` елемента.

Існують інші способи заміщення CSS. Ви пам’ятаєте рядкові стилі?

# --instructions--

Використайте рядковий стиль, щоб спробувати зробити наш `h1` елемент білим. Пам’ятайте, рядкові стилі виглядають так:

```html
<h1 style="color: green;">
```

Залиште класи `blue-text` і `pink-text` на вашому `h1` елементі.

# --hints--

Ваш `h1` елемент повинен мати клас `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Ваш `h1` елемент повинен мати клас `blue-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

Ваш `h1` елемент повинен мати ідентифікатор `orange-text`.

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'), 'orange-text');
```

Ваш `h1` елемент повинен мати рядковий стиль.

```js
assert.exists(document.querySelector('h1[style]'));
```

Ваш `h1` елемент повинен бути білим.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```
