---
id: bad87fee1348bd9aedf04756
title: Заміщення стилів у наступному CSS
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJDQug'
forumTopicId: 18253
dashedName: override-styles-in-subsequent-css
---

# --description--

Наш клас `pink-text` заміщав CSS оголошення елемента `body`!

Ми щойно довели, що наші класи заміщають CSS елемента `body`. Тож наступне логічне питання: що ми можемо зробити, щоб замістити наш клас `pink-text`?

# --instructions--

Створіть додатковий CSS клас з назвою `blue-text`, який задає елементу синій колір. Переконайтеся, що він розташований нижче за оголошення класу `pink-text`.

Застосуйте клас `blue-text` до елемента `h1` разом із класом `pink-text`, і подивимося, який з них переможе.

Застосування кількох атрибутів класу до HTML-елемента робиться через пробіл між ними так:

```html
class="class1 class2"
```

**Примітка:** Порядок, у якому класи перераховані в HTML-елементі, не має значення.

Однак порядок оголошень `class` у розділі `<style>` є важливим. Друге оголошення завжди має пріоритет над першим. Оскільки `.blue-text` оголошений другим, він заміщує атрибути `.pink-text`.

# --hints--

Ваш елемент `h1` має мати клас `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Ваш елемент `h1` має мати клас `blue-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

Обидва `blue-text` і `pink-text` мають належати одному й тому ж елементу `h1`.

```js
assert.isTrue(document.querySelector('.pink-text').classList.contains('blue-text'));
```

Ваш елемент `h1` має бути синім.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
</style>
<h1 class="pink-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }

  .blue-text {
    color: blue;
  }  
</style>
<h1 class="pink-text blue-text">Hello World!</h1>
```
