---
id: bad87fee1348bd9aedf07756
title: Заміщення всіх інших стилів за допомогою important
challengeType: 0
videoUrl: 'https://scrimba.com/c/cm24rcp'
forumTopicId: 18249
dashedName: override-all-other-styles-by-using-important
---

# --description--

Ура! Ми щойно довели, що рядкові стилі заміщають усі CSS-оголошення у вашому `style` елементі.

Але зачекайте. Існує ще один спосіб заміщення CSS. Це найпотужніший метод із усіх. Але перед тим, як його застосувати, давайте поговоримо, навіщо взагалі заміщати CSS.

У багатьох випадках ви будете використовувати CSS-бібліотеки. Вони можуть випадково заміщати ваш власний CSS. Тож коли вам абсолютно потрібно бути впевненим, що елемент має конкретний CSS, ви можете використати `!important`.

Повернемося до нашого `pink-text` оголошення класу. Пам’ятаєте, що наш `pink-text` клас був заміщений наступними оголошеннями класів, оголошеннями ідентифікаторів і рядковими стилями.

# --instructions--

Додайте ключове слово `!important` до оголошення кольору вашого елемента з рожевим текстом, щоб на 100% переконатися, що ваш `h1` елемент буде рожевим.

Приклад, як це зробити:

```css
color: red !important;
```

# --hints--

Ваш `h1` елемент має мати клас `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Ваш `h1` елемент має мати клас `blue-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

Ваш `h1` елемент має мати `id` зі значенням `orange-text`.

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'), 'orange-text');
```

Ваш `h1` елемент має мати рядковий стиль `color: white`.

```js
const commentessCode = __helpers.removeHtmlComments(code);
assert.match(commentessCode, /<h1.*style/gi);
assert.match(commentessCode, /<h1.*style.*color\s*?:/gi);
```

Ваше `pink-text` оголошення класу має містити ключове слово `!important`, щоб замістити всі інші оголошення.

```js
assert.match(__helpers.removeCssComments(code), /\.pink-text\s*?\{[\s\S]*?color:.*pink.*!important\s*;?[^\.]*\}/g);
```

Ваш `h1` елемент має бути рожевим.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color'];
assert.strictEqual(color, 'rgb(255, 192, 203)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink !important;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```
