---
id: bad87fee1348bd9aedf08726
title: Використання шістнадцяткового коду для конкретних кольорів
challengeType: 0
videoUrl: 'https://scrimba.com/c/c8W9mHM'
forumTopicId: 18350
dashedName: use-hex-code-for-specific-colors
---

# --description--

Чи знали ви, що в CSS існують інші способи представлення кольорів? Один із таких способів називається шістнадцятковим кодом, або скорочено hex-кодом.

Зазвичай ми використовуємо <dfn>десяткові</dfn> числа, або числа у десятковій системі, які використовують символи від 0 до 9 для кожної цифри. <dfn>Шістнадцяткові</dfn> (або <dfn>hex</dfn>) — це числа у системі числення з основою 16. Це означає, що вони використовують шістнадцять різних символів. Як і в десятковій системі, символи 0-9 представляють значення від нуля до дев’яти. Потім A, B, C, D, E, F представляють значення від десяти до п’ятнадцяти. Всього 0 до F можуть представляти цифру у шістнадцятковій системі, даючи нам 16 можливих значень. Більше інформації про <a href="https://www.freecodecamp.org/news/hexadecimal-number-system/" target="_blank" rel="noopener noreferrer nofollow">шістнадцяткові числа тут</a>.

У CSS ми можемо використовувати 6 шістнадцяткових цифр для представлення кольорів, по дві для червоного (R), зеленого (G) та синього (B) компонентів. Наприклад, `#000000` — це чорний колір і також найнижче можливе значення. Більше інформації про <a href="https://www.freecodecamp.org/news/rgb-color-html-and-css-guide/#whatisthergbcolormodel" target="_blank" rel="noopener noreferrer nofollow">систему кольорів RGB тут</a>.

```css
body {
  color: #000000;
}
```

# --instructions--

Замініть слово `black` у властивості колір фону елемента `body` на його шістнадцяткове представлення, `#000000`.

# --hints--

Ваш елемент `body` повинен мати `background-color` чорного кольору.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

Замість слова `black` слід використати шістнадцятковий код кольору чорного.

```js
assert.match(code, /body\s*{(([\s\S]*;\s*?)|\s*?)background.*\s*:\s*?#000(000)?((\s*})|(;[\s\S]*?}))/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: #000000;
  }
</style>
```
