---
id: bad87fee1348bd9aedf08807
title: Імпортувати вебшрифт Google
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM9MRsJ'
forumTopicId: 18200
dashedName: import-a-google-font
---

# --description--

Окрім вказування загальних шрифтів, які є на більшості операційних систем, ми також можемо вказувати нестандартні, кастомні вебшрифти для використання на нашому вебсайті. В Інтернеті є багато джерел вебшрифтів. Для цього прикладу ми зосередимося на бібліотеці Google Fonts.

Google Fonts — це безкоштовна бібліотека вебшрифтів, які ви можете використовувати у своєму CSS, посилаючись на URL шрифту.

Отже, давайте імпортуємо та застосуємо вебшрифт Google (зверніть увагу, що якщо Google заблоковано у вашій країні, вам доведеться пропустити це завдання).

Щоб імпортувати вебшрифт Google, ви можете скопіювати URL шрифту з бібліотеки Google Fonts і вставити його у свій HTML. Для цього завдання ми імпортуємо шрифт `Lobster`. Для цього скопіюйте наведений фрагмент коду і вставте його на початок свого редактора коду (перед відкриваючим елементом `style`):

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
```

Тепер ви можете використовувати шрифт `Lobster` у своєму CSS, використовуючи `Lobster` як FAMILY_NAME, як у наступному прикладі:

```css
font-family: FAMILY_NAME, GENERIC_NAME;
```

GENERIC_NAME є необов’язковим і слугує запасним варіантом шрифту на випадок, якщо інший вказаний шрифт недоступний. Це буде розглянуто у наступному завданні.

Назви сімейств шрифтів чутливі до регістру і мають бути взяті в лапки, якщо в назві є пробіл. Наприклад, для використання шрифту `"Open Sans"` потрібні лапки, а для `Lobster` — ні.

# --instructions--

Імпортуйте шрифт `Lobster` на вашу вебсторінку. Потім використайте селектор елемента, щоб встановити `Lobster` як `font-family` для вашого елемента `h2`.

# --hints--

Вам слід імпортувати шрифт `Lobster`.

```js
assert.exists(document.querySelector('link[href*="googleapis" i]'));
```

Ваш елемент `h2` має використовувати шрифт `Lobster`.

```js
const h2 = document.querySelector('h2'); 
const fontFamily = window.getComputedStyle(h2)['font-family']; 
assert.match(fontFamily, /lobster/i);
```

Використовуйте лише селектор елемента `h2`, щоб змінити шрифт.

```js
assert.match(__helpers.removeHtmlComments(code), /\s*[^\.]h2\s*\{\s*font-family\s*:\s*('|"|)Lobster\1\s*(,\s*('|"|)[a-z -]+\3\s*)?(;\s*\}|\})/gi);
```

Ваш елемент `p` має й надалі використовувати шрифт `monospace`.

```js
const paragraphElement = document.querySelector('p');
const fontFamily = window.getComputedStyle(paragraphElement)['font-family']; 
assert.match(fontFamily, /monospace/i);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: red;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
<style>
  .red-text {
    color: red;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }

  h2 {
    font-family: Lobster;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
