---
id: bad87fee1348bd9aedf08805
title: Використання селекторів CSS для стилізації елементів
challengeType: 0
videoUrl: 'https://scrimba.com/c/cJKMBT2'
forumTopicId: 18349
dashedName: use-css-selectors-to-style-elements
---

# --description--

За допомогою CSS існує сотні CSS-властивостей, які ви можете використати, щоб змінити вигляд елемента на вашій сторінці.

Коли ви додали `<h2 style="color: red;">CatPhotoApp</h2>`, ви стилізували цей окремий `h2` елемент за допомогою рядкового CSS, що означає каскадні таблиці стилів.

Це один зі способів задати стиль елемента, але є кращий спосіб застосувати CSS.

У верхній частині вашого коду створіть `style` блок ось так:

```html
<style>
</style>
```

Всередині цього блоку стилів ви можете створити <dfn>селектор CSS</dfn> для всіх `h2` елементів. Наприклад, якщо ви хочете, щоб усі `h2` елементи були синіми, ви додасте правило стилю, яке виглядає так:

```html
<style>
  h2 {
    color: red;
  }
</style>
```

Зверніть увагу, що важливо мати як відкриваючі, так і кінцеві фігурні дужки (`{` і `}`) навколо кожного правила стилю елемента. Також потрібно переконатися, що визначення стилю вашого елемента знаходиться між відкриваючим і кінцевим тегами style. Нарешті, обов’язково додайте крапку з комою в кінці кожного правила стилю вашого елемента.

# --instructions--

Видаліть атрибут стилю вашого `h2` елемента і замість цього створіть CSS `style` блок. Додайте необхідний CSS, щоб усі `h2` елементи стали синіми.

# --hints--

Атрибут `style` слід видалити з вашого `h2` елемента.

```js
assert.notExists(document.querySelector('h2').getAttribute('style'));
```

Ви маєте створити `style` елемент.

```js
assert.exists(document.querySelector('style:not(.fcc-hide-header)'));
assert.isAtLeast(document.querySelectorAll('style:not(.fcc-hide-header)').length, 1);
```

Ваш `h2` елемент має бути синім.

```js
const h2Element = document.querySelector('h2');
const color = window.getComputedStyle(h2Element)['color']; 
assert.strictEqual(color, 'rgb(0, 0, 255)');
```

Ваше оголошення `h2` у таблиці стилів має бути дійсним із крапкою з комою та кінцевою дужкою.

```js
assert.match(__helpers.removeCssComments(code), /h2\s*\{\s*color\s*:.*;\s*\}/g);
```

Всі ваші `style` елементи мають бути дійсними і мати кінцеві теги.

```js
assert.match(__helpers.removeHtmlComments(code), /<\/style>/g);

const closingElementLength = __helpers.removeHtmlComments(code).match(/<\/style>/g).length;
const openingElementsLength = __helpers.removeHtmlComments(code).match(/<style((\s)*((type|media|scoped|title|disabled)="[^"]*")?(\s)*)*>/g).length;
assert.strictEqual(closingElementLength, openingElementsLength);
```

# --seed--

## --seed-contents--

```html
<h2 style="color: red;">CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<style>
  h2 {
    color: blue;
  }
</style>
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
