---
id: bad87fee1348bd9aedf08808
title: "Вкажіть, як шрифти мають деградувати"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cpVKBfQ'
forumTopicId: 18304
dashedName: specify-how-fonts-should-degrade
---

# --description--

Існує кілька шрифтів за замовчуванням, які доступні у всіх браузерах. Ці загальні сімейства шрифтів включають `monospace`, `serif` та `sans-serif`.

Коли один шрифт недоступний, ви можете вказати браузеру "деградувати" до іншого шрифту.

Наприклад, якщо ви хочете, щоб елемент використовував шрифт `Helvetica`, але деградував до шрифту `sans-serif`, коли `Helvetica` недоступний, ви вкажете це так:

```css
p {
  font-family: Helvetica, sans-serif;
}
```

Назви загальних сімейств шрифтів не чутливі до регістру. Також їм не потрібні лапки, бо це ключові слова CSS.

# --instructions--

Щоб почати, застосуйте шрифт `monospace` до елемента `h2`, щоб тепер він мав два шрифти — `Lobster` і `monospace`.

У останньому завданні ви імпортували шрифт `Lobster` за допомогою тега `link`. Тепер перетворіть цей імпорт шрифту `Lobster` з Google Fonts на коментар (використовуючи HTML-коментарі, які ви вивчали раніше), щоб він більше не був доступний. Зверніть увагу, як ваш елемент `h2` деградує до шрифту `monospace`.

**Примітка:** Якщо у вас встановлений шрифт `Lobster` на комп’ютері, ви не побачите деградації, бо браузер зможе знайти цей шрифт.

# --hints--

Ваш елемент h2 має використовувати шрифт `Lobster`.

```js
const h2Element = document.querySelector('h2');
const fontFamily = window.getComputedStyle(h2Element)['font-family']; 
assert.match(fontFamily, /^"?lobster/i);
```

Ваш елемент h2 має деградувати до шрифту `monospace`, коли `Lobster` недоступний.

```js
assert.match(__helpers.removeCssComments(code), /\s*h2\s*\{\s*font-family\s*\:\s*(\'|"|)Lobster\1\s*,\s*monospace\s*;?\s*\}/gi);
```

Ви маєте перетворити на коментар виклик Google для шрифту `Lobster`, додавши `<!--` перед ним.

```js
assert.match(code, /<!--[^fc]/gi);
```

Ви маєте закрити коментар, додавши `-->`.

```js
assert.match(code, /[^fc]-->/gi);
```

# --seed--

## --seed-contents--

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
<style>
  .red-text {
    color: red;
  }

  h2 {
    font-family: Lobster;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<!--<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">-->
<style>
  .red-text {
    color: red;
  }

  h2 {
    font-family: Lobster, monospace;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>
  
  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
