---
id: bad87fee1348bd9aedf08756
title: Віддавати перевагу одному стилю над іншим
challengeType: 0
videoUrl: 'https://scrimba.com/c/cZ8wnHv'
forumTopicId: 18258
dashedName: prioritize-one-style-over-another
---

# --description--

Іноді ваші HTML-елементи отримують кілька стилів, які суперечать один одному.

Наприклад, ваш `h1` елемент не може бути одночасно і зеленим, і рожевим.

Давайте подивимося, що станеться, якщо ми створимо клас, який робить текст рожевим, а потім застосуємо його до елемента. Чи замістить наш клас `body` властивість `color: green;` CSS елемента?

# --instructions--

Створіть CSS-клас з назвою `pink-text`, який задає елементу рожевий колір.

Додайте вашому `h1` елементу клас `pink-text`.

# --hints--

Ваш `h1` елемент має мати клас `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Ваш `<style>` має мати CSS-клас `pink-text`, який змінює `color`.

```js
assert.match(__helpers.removeCssComments(code), /\.pink-text\s*\{\s*color\s*:\s*.+\s*;?\s*\}/g);
```

Ваш `h1` елемент має бути рожевим.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 192, 203)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
</style>
<h1>Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
</style>
<h1 class="pink-text">Hello World!</h1>
```
