---
id: 66bb3e20d3dc5b6d0a21f5dd
title: 构建一个实时计数器
challengeType: 25
dashedName: build-a-real-time-counter
demoType: onClick
---

# --description--

在本实验中，你将使用 JavaScript 构建一个实时字符计数器。字符计数器将显示在 `textarea` 元素中输入的字符数。计数器将在用户在 `textarea` 中输入时实时更新。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个 `textarea` 元素，其 `id` 为 `text-input`。
2. 应该有一个 `p` 元素，其 `id` 为 `char-count`。它最初应包含文本 `Character Count: 0/50`。此占位符文本应被 `textarea` 元素中当前的字符计数所替换。
3. 当 `#text-input` 元素包含文本 `hello coder` 时，`#char-count` 元素应包含文本 `"Character Count: 11/50"`
4. 当 `#text-input` 元素包含文本 `i am a programmer` 时，`#char-count` 元素应包含文本 `"Character Count: 17/50"`
5. 当 `#text-input` 元素包含文本 `hello world` 时，`#char-count` 元素应包含文本 `"Character Count: 11/50"`
6. 当 `#text-input` 元素包含文本 `I am learning a new language and it is called c++.` 时，`#char-count` 元素应包含文本 `"Character Count: 50/50"`。**注意**：虽然 `maxlength` 属性可以实现正确的功能性，但本实验要求仅使用 JavaScript 方案。
7. 用户不应能够在 `textarea` 元素中输入超过 `50` 个字符。当字符数达到 `50` 时，任何多余的输入都应自动修剪，并且应以 `red` 显示文本 `Character Count: 50/50`。

**注意：** 请确保在你的超文本标记语言中链接你的样式表和 JavaScript 文件。

# --hints--

你应该有一个 `textarea` 元素，其 `id` 为 `text-input`。 

```js
const textInput = document.querySelector('textarea');
assert.equal(textInput.id, 'text-input');
```

你应该有一个 `p` 元素，其 `id` 为 `char-count`。

```js
assert.exists("p#char-count");
```

`#char-count` 元素应最初包含文本 `Character Count: 0/50`。

```js
const charCount = document.querySelector('p#char-count')?.innerText;
assert.equal(charCount, 'Character Count: 0/50');
```

当 `#text-input` 元素包含文本 `hello coder` 时，`#char-count` 元素应包含文本 `"Character Count: 11/50"`

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "hello coder";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(charCount.innerText.trim(), 'Character Count: 11/50');
```

当字符数为 `50` 时，文本应显示为 `red`。

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "I am learning a new language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
const computedStyle = window.getComputedStyle(charCount);
assert.oneOf(computedStyle.color, [
  "red",
  "rgb(255, 0, 0)",
  "rgb(100%, 0%, 0%)",
  "rgba(255, 0, 0, 1)",
  "rgba(100%, 0%, 0%, 1)",
  "hsl(0, 100%, 50%)",
  "hsla(0, 100%, 50%, 1)"
]);
```

如果字符数大于或相等于 `50`，多余的输入应被裁剪。

```js
const textInput = document.getElementById('text-input');
textInput.value = "I am learning a new programming language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(textInput.value.length, 50);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real Time Counter</title>
</head>

<body>

</body>


</html>
```

```css

```

```js
    
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real-time Character Counter</title>
    <link rel="stylesheet" href="styles.css">

<body>
    <div class="container">
        <h1>Real-Time Character Counter</h1>
        <textarea id="text-input" placeholder="Type something..."></textarea>
        <p id="char-count">Character Count: 0/50</p>
    </div>

<script src="script.js"></script>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    font-family: Arial, sans-serif;
    background-color: #c0c1ec;
    margin: 0;
}

h1{
    background-color: #f2f48b;
    padding-left: 10px;
    padding-right: 10px;
}
.container {
    text-align: center;
}

textarea {
    width: 300px;
    height: 100px;
    font-size: 18px;
    padding: 10px;
    border: 1px solid #ccc;
    border-radius: 5px;
    resize: none;
    display: block;
    margin: 0 auto;
    border: 1px solid slategray;
    box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

textarea:focus {
    outline: none;
    border-color: #f2f48b;
}

#char-count {
    margin-top: 10px;
    font-size: 18px;
    color: darkslategray;
    font-weight: bold;
}
```

```js
document.getElementById("text-input").addEventListener("input", function () {
  let textInput = document.getElementById("text-input");
  let charCount = textInput.value.length;

  if (charCount > 50) {
    textInput.value = textInput.value.substring(0, 50);
    charCount = 50;
  }

  let charCountDisplay = document.getElementById("char-count");
  charCountDisplay.innerHTML = `Character Count: ${charCount}/50`;

  if (charCount === 50) {
    charCountDisplay.style.color = "red";
  } 

});
```
