---
id: 698c58745b73289bb6c89ece
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

L’ultima canzone si intitola `Scratching the Surface`, l’artista è Quincy Larson e il file da usare è `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3`.

Aggiungi l’ultima canzone per completare il lettore musicale seguendo lo stesso pattern usato per le prime due canzoni.

# --hints--

Dovresti avere un terzo elemento `h2`.

```js
assert.exists(document.querySelector('h2:nth-of-type(3)'));
```

Il tuo terzo elemento `h2` dovrebbe contenere il testo `Scratching the Surface`.

```js
const h2 = document.querySelector('h2:nth-of-type(3)');
assert.strictEqual(h2?.textContent.trim(), 'Scratching the Surface');
```

Dovresti avere un elemento `p`.

```js
assert.exists(document.querySelector('p:nth-of-type(3)'));
```

Il tuo elemento `p` dovrebbe contenere il testo `Artist: Quincy Larson`.

```js
const p = document.querySelector('p:nth-of-type(3)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

Dovresti avere un terzo elemento `audio`.

```js
assert.exists(document.querySelector('audio:nth-of-type(3)'));
```

Il nuovo elemento `audio` dovrebbe essere sotto l’elemento `p`.

```js
assert.exists(document.querySelector('p:nth-of-type(3) + audio:nth-of-type(3)'));
```

Il tuo elemento `audio` dovrebbe avere un attributo `src`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('src'));
```

Il tuo elemento `audio` dovrebbe avere un attributo `src` con valore `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3");
```

Il tuo elemento `audio` dovrebbe avere l’attributo `loop`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('loop'));
```

Il tuo elemento `audio` dovrebbe avere l’attributo `controls`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

  <h2>Scratching the Surface</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3" loop controls></audio>
  
</body>
</html>
```
