---
id: 698c55e7aaacfae48a227459
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

Aggiungi un nuovo elemento `h2` sotto l'elemento `audio` con il testo `Cruising for a Musing`.

Sotto `h2`, aggiungi un elemento `p` con il testo `Artist: Quincy Larson`.

# --hints--

Dovresti avere un secondo elemento `h2`.

```js
assert.exists(document.querySelector('h2:nth-of-type(2)'));
```

Il tuo secondo elemento `h2` dovrebbe contenere il testo `Cruising for a Musing`.

```js
const h2 = document.querySelector('h2:nth-of-type(2)');
assert.strictEqual(h2?.textContent.trim(), 'Cruising for a Musing');
```

Dovresti avere un secondo elemento `p`.

```js
assert.exists(document.querySelector('p:nth-of-type(2)'));
```

Il tuo secondo elemento `p` dovrebbe contenere il testo `Artist: Quincy Larson`.

```js
const p = document.querySelector('p:nth-of-type(2)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
