---
id: 698c5710136539797f296098
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Sotto l'elemento `p`, aggiungi un elemento `audio`, dagli un attributo `src` con valore `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3` e gli attributi `loop` e `controls`.

# --hints--

Dovresti avere un secondo elemento `audio`.

```js
assert.exists(document.querySelector('audio:nth-of-type(2)'));
```

Il nuovo elemento `audio` dovrebbe essere sotto l'elemento `p`.

```js
assert.exists(document.querySelector('p:nth-of-type(2) + audio:nth-of-type(2)'));
```

Il tuo elemento `audio` dovrebbe avere un attributo `src`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('src'));
```

Il tuo elemento `audio` dovrebbe avere un attributo `src` con valore `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3");
```

Il tuo elemento `audio` dovrebbe avere l'attributo `loop`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('loop'));
```

Il tuo elemento `audio` dovrebbe avere l'attributo `controls`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
