---
id: 698c4f4e965bd853ad5084e1
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Sotto il `h1`, aggiungi un elemento `h2` con il testo `Can't Stay Down`, che è il titolo della prima canzone.

Sotto l'elemento `h2` aggiungi un elemento `p` con il testo `Artist: Quincy Larson`.

# --hints--

Dovresti avere un elemento `h2`.

```js
assert.exists(document.querySelector('h2'));
```

Il tuo elemento `h2` dovrebbe essere sotto l'elemento `h1`.

```js
assert.exists(document.querySelector('h1 + h2'));
```

Il tuo elemento `h2` dovrebbe contenere il testo `Can't Stay Down`.

```js
const h2 = document.querySelector('h2');
assert.strictEqual(h2?.textContent.trim(), "Can't Stay Down");
```

Dovresti avere un elemento `p`.

```js
assert.exists(document.querySelector('p'));
```

Il tuo elemento `p` dovrebbe contenere il testo `Artist: Quincy Larson`.

```js
const p = document.querySelector('p');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
