---
id: 698c5229e47d57c6dfaa91bd
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

L'attributo `controls` fornisce i controlli di riproduzione inclusi pausa/ripresa, ricerca e controllo del volume per l'elemento `audio`.

L'attributo `controls` è un attributo booleano e non necessita di un valore.

Aggiungi l'attributo `controls` all'elemento `audio`.

Ora dovresti vedere l'elemento `audio` visualizzato sulla pagina.

# --hints--

Il tuo elemento `audio` dovrebbe avere l'attributo `controls`.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop></audio>
--fcc-editable-region--

</body>
</html>
```
