---
id: 698c50535b54b2d75fa985fb
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

L'attributo `loop` farà ripartire l'audio una volta terminata la riproduzione. Pensa a un meme di internet che ripete la riproduzione. Omettere l'attributo `loop` farà sì che l'audio venga riprodotto una sola volta.

L'attributo `loop` è un attributo booleano e non necessita di un valore.

Aggiungi l'attributo `loop` all'elemento `audio`.

# --hints--

Il tuo elemento `audio` dovrebbe avere l'attributo `loop`.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3"></audio>
--fcc-editable-region--

</body>
</html>
```
