---
id: ab6137d4e35944e21037b769
title: بناء محول إلى حالة العنوان
challengeType: 26
dashedName: build-a-title-case-converter
---

# --description--

في هذا المختبر ستنشئ دالة تحول سلسلة نصية إلى حالة العنوان. حالة العنوان تعني أن الحرف الأول من كل كلمة يكون بحرف كبير وبقية الكلمة بحروف صغيرة.

`"Web Development Is Awesome"` مثال على سلسلة نصية بحالة العنوان.

**الهدف**: حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تملك دالة `titleCase` تأخذ سلسلة نصية كمعلمة.
1. يجب أن تُرجع دالة `titleCase` سلسلة نصية يكون فيها الحرف الأول من كل كلمة بحرف كبير وبقية الكلمة بحروف صغيرة.
1. يجب أن تُرجع `titleCase("I like to code")` السلسلة `"I Like To Code"`.
1. يجب أن تُرجع `titleCase("javaScript is fun")` السلسلة `"Javascript Is Fun"`.

# --hints--

يجب أن تنشئ دالة باسم `titleCase`.

```js
assert.isFunction(titleCase);
```

يجب أن تأخذ `titleCase` معلمة واحدة.

```js
assert.lengthOf(titleCase, 1);
```


يجب أن تُرجع `titleCase("I'm a little tea pot")` سلسلة نصية.

```js
assert.isString(titleCase("I'm a little tea pot"));
```

يجب أن تُرجع `titleCase("I'm a little tea pot")` السلسلة `I'm A Little Tea Pot`.

```js
assert.strictEqual(titleCase("I'm a little tea pot"), "I'm A Little Tea Pot");
```

يجب أن تُرجع `titleCase("sHoRt AnD sToUt")` السلسلة `Short And Stout`.

```js
assert.strictEqual(titleCase('sHoRt AnD sToUt'), 'Short And Stout');
```

يجب أن تُرجع `titleCase("HERE IS MY HANDLE HERE IS MY SPOUT")` السلسلة `Here Is My Handle Here Is My Spout`.

```js
assert.strictEqual(
  titleCase('HERE IS MY HANDLE HERE IS MY SPOUT'),
  'Here Is My Handle Here Is My Spout'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function titleCase(str) {
  return str
    .split(' ')
    .map(word => word.charAt(0).toUpperCase() + word.substring(1).toLowerCase())
    .join(' ');
}
```
