---
id: 69de4308e652aaeb5ccdb03e
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este taller, practicarás el uso de los métodos de cadena de Python.

Comienza creando una variable llamada `email` y asigna la cadena `alice.johnson@company.com`.

# --hints--

Debes crear una variable llamada `email`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email')
`) })
```

Debes asignar la cadena `alice.johnson@company.com` a `email`. Recuerda encerrar el texto entre comillas simples o comillas dobles.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('email').is_equivalent("email = 'alice.johnson@company.com'")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
