---
id: 69e20876baed33de636b989a
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

A veces, puede que quieras encontrar la posición de un carácter específico dentro de una cadena. Python proporciona el método integrado `.find()` para este propósito. Devuelve un entero que representa el índice (posición) de la primera aparición de la subcadena dada.

Aquí tienes un ejemplo:

```py
text = 'hello'
print(text.find('e'))  # 1
```

Ahora, llama al método `.find()` en `email` y pásale `'@'` para encontrar la posición del signo arroba en la cadena `email`. Asigna el resultado a una variable llamada `at_position`.

Luego imprime `at_position` para que la salida sea: `Position of @: 13`.

# --hints--

Debes crear una variable llamada `at_position`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_position')
`) })
```

Debes asignar a `at_position` el valor `email.find('@')`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_position').is_equivalent("at_position = email.find('@')")
`) })
```

Debes imprimir `at_position` para que la salida sea `Position of @: 13`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Position of @: {at_position}')",
    "print('Position of @: ' + str(at_position))",
    "print('Position of @:', at_position)",
    "print('Position of @:', str(at_position))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')
--fcc-editable-region--

--fcc-editable-region--
```
