---
id: 69e20a0451c7a3ef189514d3
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Ahora que sabes la posición de `@` en el correo electrónico, puedes extraer el nombre de usuario, que está antes del signo `@`.

En el taller anterior, extrajiste una porción de una cadena usando slicing. Como recordatorio, la sintaxis es:

```python
text = 'hello world'
print(text[:5])  # hello
```

Usa slicing en `email` para extraer todo lo que está antes de `at_position` y guarda el resultado en una variable llamada `username`.

Luego imprime `username` para que la salida sea `Username: alice.johnson`.

# --hints--

Debes crear una variable llamada `username`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('username')
`) })
```

Debes usar slicing con `at_position` en `email` para asignar un valor a `username`.

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('username')
assert node.is_equivalent("username = email[:at_position]") or node.is_equivalent("username = email[0:at_position]")
`) })
```

Debes imprimir `username` para que la salida sea `Username: alice.johnson`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {username}')",
    "print('Username: ' + username)",
    "print('Username:', username)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

--fcc-editable-region--

--fcc-editable-region--
```
