---
id: 69e6299710319e7cb8ebdc37
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Los datos del mundo real suelen estar desordenados, con espacios extra y mayúsculas y minúsculas inconsistentes.

Crea una variable `messy_name` y asígnale la cadena `sARaH dAVis` con dos espacios al principio y dos espacios al final. Luego imprímela para que la salida sea:

```md
Original name: '  sARaH dAVis  '
```

Lo limpiarás en los próximos pasos.

# --hints--

Debes crear una variable llamada `messy_name`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_name')
`) })
```

Debes asignar <code>'&nbsp;&nbsp;sARaH dAVis&nbsp;&nbsp;'</code> a `messy_name`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_name').is_equivalent("messy_name = '  sARaH dAVis  '")
`) })
```

Debes imprimir `messy_name` para que la salida sea `Original name: '  sARaH dAVis  '`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original name: '{messy_name}'\\")",
    "print(\\"Original name: '\\" + messy_name + \\"'\\")",
    "print(\\"Original name: '\\", messy_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

--fcc-editable-region--

--fcc-editable-region--
```
