---
id: 69e62ce96df237884da25282
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

El método `.lower()` convierte todos los caracteres de una cadena a minúsculas. Esto es útil para normalizar direcciones de correo electrónico:

```python
upperGreet = 'HELLO WORLD'
lowerGreet = upperGreet.lower()
print(lowerGreet)  # Output: hello world
```

Puedes llamar a `.strip()` y `.lower()` en una cadena de esta manera:

```python
messy = '  HELLO WORLD  '
clean = messy.strip().lower()
print(clean)  # Output: hello world
```

Crea una variable `messy_email` y asígnale la cadena `SARAH.DAVIS@COMPANY.COM` con dos espacios al inicio y dos espacios al final.

Además, llama a `.strip().lower()` sobre `messy_email` y asigna el resultado a una variable `clean_email`.

Luego imprime ambas para que la salida sea:

```md
Original email: '  SARAH.DAVIS@COMPANY.COM  '
Cleaned email: 'sarah.davis@company.com'
```

# --hints--

Debes crear una variable llamada `messy_email`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_email')
`) })
```

Debes asignar <code>'&nbsp;&nbsp;SARAH.DAVIS@COMPANY.COM&nbsp;&nbsp;'</code> a `messy_email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_email').is_equivalent("messy_email = '  SARAH.DAVIS@COMPANY.COM  '")
`) })
```

Debes crear una variable llamada `clean_email`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('clean_email')
`) })
```

Debes llamar a `.strip().lower()` sobre `messy_email` y asignar el resultado a una variable `clean_email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_email').is_equivalent("clean_email = messy_email.strip().lower()")
`) })
```

Debes imprimir `messy_email` para que la salida sea `Original email: '  SARAH.DAVIS@COMPANY.COM  '`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original email: '{messy_email}'\\")",
    "print(\\"Original email: '\\" + messy_email + \\"'\\")",
    "print(\\"Original email: '\\", messy_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

Debes imprimir `clean_email` para que la salida sea `Cleaned email: 'sarah.davis@company.com'`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned email: '{clean_email}'\\")",
    "print(\\"Cleaned email: '\\" + clean_email + \\"'\\")",
    "print(\\"Cleaned email: '\\", clean_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
