---
id: 69e6284f3b7659c5a05c2a79
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Puedes verificar si una subcadena existe dentro de una cadena usando la palabra clave `in`. Devuelve `True` si la subcadena existe y `False` si no. Aquí tienes un ejemplo:

```py
sentence = 'I love Python'
print('Python' in sentence) # True
print('Java' in sentence) # False
```

Crea una variable llamada `has_at` que verifique si el carácter `@` está en `email`.

Luego imprime `has_at` para que la salida sea: `Has @ symbol: True`.

# --hints--

Debes crear una variable llamada `has_at`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_at')
`) })
```

Tu variable `has_at` debe tener el valor `'@' in email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_at').is_equivalent("has_at = '@' in email")
`) })
```

Debes imprimir `has_at` para que la salida sea `Has @ symbol: True`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has @ symbol: {has_at}')",
    "print('Has @ symbol: ' + str(has_at))",
    "print('Has @ symbol:', has_at)",
    "print('Has @ symbol:', str(has_at))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

--fcc-editable-region--

--fcc-editable-region--
```
