---
id: 69e73a141addc6f442924b0d
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

Ya usaste `.find()` y slicing en la variable `email`. Ahora usarás la misma técnica en `clean_email`.

Encuentra la posición del carácter `@` en `clean_email` y guárdala en `at_pos`. Luego usa slicing para extraer todo lo que está antes de `@` y guárdalo en `email_user`.

Luego imprime `clean_email` y `email_user` para que la salida sea:

```md
Email: sarah.davis@company.com
Username: sarah.davis
```

# --hints--

Debes crear una variable llamada `at_pos`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_pos')
`) })
```

Debes llamar a `.find('@')` en `clean_email` y asignar el resultado a la variable `at_pos`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_pos').is_equivalent("at_pos = clean_email.find('@')")
`) })
```

Debes crear una variable llamada `email_user`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email_user')
`) })
```

Debes usar slicing con `at_pos` en `clean_email` y asignar el resultado a la variable `email_user`.

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('email_user')
assert node.is_equivalent("email_user = clean_email[:at_pos]") or node.is_equivalent("email_user = clean_email[0:at_pos]")
`) })
```

Debes imprimir `clean_email` para que la salida sea `Email: sarah.davis@company.com`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {clean_email}')",
    "print('Email: ' + clean_email)",
    "print('Email:', clean_email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

Debes imprimir `email_user` para que la salida sea `Username: sarah.davis`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {email_user}')",
    "print('Username: ' + email_user)",
    "print('Username:', email_user)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

--fcc-editable-region--

--fcc-editable-region--
```
