---
id: 6a71b3e49668c3be461257b5
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

El método `.title()` pone en mayúscula la primera letra de cada palabra. Aquí tienes un ejemplo:

```py
lowercase = 'hello world'
title_case = lowercase.title()
print(title_case)  # Output: Hello World
```

Crea una variable llamada `clean_name` y asígnale el resultado de llamar a `.title()` sobre `messy_name`.

# --hints--

Debes crear una variable llamada `clean_name`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('clean_name')
`) })
```

Debes usar `.title()` en `messy_name` para asignar un valor a `clean_name`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_name').is_equivalent("clean_name = messy_name.title()")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
