---
id: 6a86b9e64da746838e67fb13
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

El método `.strip()` elimina los espacios en blanco de ambos extremos de una cadena.

```python
messy = '  hello world  '
clean = messy.strip()
print(clean)  # Output: hello world
```

Puedes encadenar `.strip()` y `.title()` para eliminar los espacios en blanco de ambos extremos de una cadena y capitalizar la primera letra de cada palabra:

```python
messy = '  hello world  '
clean = messy.strip().title()
print(clean)  # Output: Hello World
```

Actualiza el valor de `clean_name` para llamar a `.strip()` sobre `messy_name` antes de `.title()`.

Después, imprime `clean_name` para que la salida sea `Cleaned name: 'Sarah Davis'`.

# --hints--

Debes actualizar `clean_name` para llamar a `.strip()` sobre `messy_name` antes de `.title()`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_name').is_equivalent("clean_name = messy_name.strip().title()")
`) })
```

Debes imprimir `clean_name` para que la salida sea `Cleaned name: 'Sarah Davis'`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned name: '{clean_name}'\\")",
    "print(\\"Cleaned name: '\\" + clean_name + \\"'\\")",
    "print(\\"Cleaned name: '\\", clean_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--
clean_name = messy_name.title()
--fcc-editable-region--
```
