---
id: 6a870cc60c23f0332a2fe788
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Recuerda de un taller anterior que puedes usar una f-string para incluir una variable dentro de una cadena, escribiendo una `f` antes de la comilla de apertura y envolviendo la variable en `{}`.

Aquí tienes un recordatorio de cómo funciona:

```python
name = 'Alice'
print(f'Hello, {name}')  # Hello, Alice
```

Ahora, imprime `email` para que la salida sea: `Email: alice.johnson@company.com`.

# --hints--

Debes imprimir `email` para que la salida sea `Email: alice.johnson@company.com`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {email}')",
    "print('Email: ' + email)",
    "print('Email:', email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
--fcc-editable-region--

--fcc-editable-region--
```
