---
id: 673339d96c56d36e963e888e
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、JavaScriptのMapとSetオブジェクトを練習しながら、植物園のカタログを作成します。

カタログ内の各植物を表すために、3つのプロパティを持つオブジェクトを使います。

- `commonName`：植物の一般名
- `scientificName`：植物の学名（またはラテン名）
- `cultivar`：植物の品種

まず、`Spanish lavender`（学名`Lavandula stoechas`）という品種`Ballerina`のために、上記3つのプロパティを持つオブジェクトを作成し、`ballerina`という名前の変数に代入してください。

# --hints--

`ballerina`という名前の変数があるはずです。

```js
assert.exists(ballerina);
```

`ballerina`変数にオブジェクトを代入してください。

```js
assert.isObject(ballerina);
```

`ballerina`オブジェクトは`commonName`プロパティを`"Spanish lavender"`に設定してください。

```js
assert.strictEqual(ballerina.commonName, "Spanish lavender");
```

`ballerina`オブジェクトは`scientificName`プロパティを`"Lavandula stoechas"`に設定してください。

```js
assert.strictEqual(ballerina.scientificName, "Lavandula stoechas");
```

`ballerina`オブジェクトは`cultivar`プロパティを`"Ballerina"`に設定してください。

```js
assert.strictEqual(ballerina.cultivar, "Ballerina");
```

`ballerina`オブジェクトは3つのキーを持つはずです。

```js
assert.lengthOf(Object.keys(ballerina), 3);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
