---
id: 6733b5e6f210c1bd1afdf8b3
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

残りの植物オブジェクトはすでに追加されています。次はカタログを作成しましょう。

通常のJavaScriptオブジェクトと同様に、`Map`オブジェクトはキーと値のペアの集合を含みます。`Map`オブジェクトを作成するには、`Map`コンストラクタを使うことができます。

```js
const myMap = new Map();
```

`catalog`変数を宣言し、`Map`オブジェクトを代入してください。

# --hints--

`catalog`という名前の変数があるはずです。

```js
assert.exists(catalog);
```

`Map`変数に`catalog`オブジェクトを代入してください。

```js
assert.instanceOf(catalog, Map);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--

--fcc-editable-region--
```
