---
id: 6733b5632fb854bc6c58dc0d
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

同じ3つのプロパティを持つ別のオブジェクトを作成し、スペインラベンダーの別の品種である`Pretty Polly`を表し、`prettyPolly`という変数に代入してください。

# --hints--

`prettyPolly`という名前の変数があるはずです。

```js
assert.exists(prettyPolly);
```

`prettyPolly`変数にオブジェクトを代入する必要があります。

```js
assert.isObject(prettyPolly);
```

`prettyPolly`オブジェクトには、`commonName`プロパティが`"Spanish lavender"`に設定されている必要があります。

```js
assert.strictEqual(prettyPolly.commonName, "Spanish lavender");
```

`prettyPolly`オブジェクトには、`scientificName`プロパティが`"Lavandula stoechas"`に設定されている必要があります。

```js
assert.strictEqual(prettyPolly.scientificName, "Lavandula stoechas");
```

`prettyPolly`オブジェクトには、`cultivar`プロパティが`"Pretty Polly"`に設定されている必要があります。

```js
assert.strictEqual(prettyPolly.cultivar, "Pretty Polly");
```

`prettyPolly`オブジェクトには3つのキーがある必要があります。

```js
assert.lengthOf(Object.keys(prettyPolly), 3);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}
--fcc-editable-region--

--fcc-editable-region--
```
