---
id: 6734c10a327760a665b7d5b9
title: ステップ 13
challengeType: 1
dashedName: step-13
---

# --description--

次の数ステップでは、植物を販売したときにカタログを更新できる関数に取り組みます。

3つの引数を受け取る`sellPlants`という名前の関数を作成してください。引数は順に、植物オブジェクト、鉢のサイズ、販売する鉢の数です。

指定された鉢の数が利用可能な鉢の数より多い場合、関数は`Not enough <pot-size> size pots for <item-name>. Only <pot-number> left.`を返すようにしてください。`<pot-size>`は鉢のサイズに、`<pot-number>`はそのサイズの残りの鉢の数に、`<item-name>`は植物の学名と、その後に半角スペースを挟んで単一引用符で囲まれた栽培品種名に置き換えてください。

すべてが正しく動作するか確認するために、`sellPlants`関数を`ballerina`オブジェクト、`"small"`、`25`の引数で呼び出した結果をコンソールにログ出力してください。

# --hints--

`catalog`マップは変更しないでください。

```js
const testCatalog = new Map();
testCatalog.set(ballerina, { small: 20, medium: 15, large: 12 });
testCatalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
testCatalog.set(willowVale, { small: 3, medium: 5, large: 0 });
testCatalog.set(hidcote, { small: 33, medium: 13, large: 18 });
testCatalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
testCatalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
assert.deepEqual(catalog, testCatalog);
```

`sellPlants`という名前の関数があるはずです。

```js
assert.isFunction(sellPlants);
```

`sellPlants`関数は3つの引数を受け取るべきです。

```js
assert.lengthOf(sellPlants, 3);
```

`sellPlants(ballerina, "small", 25)`は`Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.`を返すべきです。

```js
assert.equal(sellPlants(ballerina, "small", 25), "Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.")

// prevent hardcoding
assert.equal(sellPlants(hidcote, "large", 95), "Not enough large size pots for Lavandula angustifolia 'Hidcote'. Only 18 left.")

assert.equal(sellPlants(royalCrown, "medium", 95), "Not enough medium size pots for Lavandula dentata 'Royal Crown'. Only 22 left.")
```

`sellPlants(ballerina, "small", 25)`をコンソールにログ出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*sellPlants\s*\(\s*ballerina\s*,\s*("|')small\1\s*,\s*25\s*\)\s*\)/);
```

`sellPlants(ballerina, "small", 10)`は`undefined`を返すべきです。

```js
assert.isUndefined(sellPlants(ballerina, "small", 10));
// prevent hardcoding
assert.isUndefined(sellPlants(hidcote, "large", 5));
assert.isUndefined(sellPlants(royalCrown, "medium", 5));
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog)
```
