---
id: 6734f0ae9d1cf4dfd43a7805
title: ステップ 15
challengeType: 1
dashedName: step-15
---

# --description--

mapオブジェクトから要素を削除するには、キーを引数に取り、そのキーに関連付けられたエントリーを削除する`delete`メソッドを使えます。このメソッドは、削除が成功したかどうかを示すブール値を返します。

`removePlant`という名前の関数を作成し、植物オブジェクトを引数に取り、カタログから削除してください。指定された植物がカタログから削除されたかどうかに応じて、`true`または`false`を返すようにしてください。

その後、文字列`"ballerina"`を引数にして新しい関数を呼び出した結果をコンソールに出力してください。

# --hints--

`removePlant`という名前の関数を作成してください。

```js
assert.isFunction(removePlant);
```

`removePlant`関数は1つの引数を取るべきです。

```js
assert.lengthOf(removePlant, 1);
```

`removePlant`関数は`delete`の`catalog`メソッドを呼び出すべきです。

```js
let flag = false;
const temp = catalog.delete;
try {
    catalog.delete = (p) => {flag = true}
    removePlant("");
    assert.isTrue(flag);
} finally {
    catalog.delete = temp;
}
```

`removePlant(item)`はカタログから`item`を削除するべきです。

```js
const testArr = [];
const temp = catalog.delete;
try {
    catalog.delete = (p) => {testArr.push(p)}
    removePlant("test");
    assert.deepEqual(testArr, ["test"]);
} finally {
    catalog.delete = temp;
}
```

`removePlant`関数は削除が成功したかどうかを示すブール値を返すべきです。

```js
try {
    assert.isFalse(removePlant("test"));
    assert.isTrue(removePlant(royalCrown));
} finally {
    catalog.set(royalCrown, { small: 40, medium: 22, large: 9 })
}
```

`removePlant("ballerina")`をコンソールに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*removePlant\s*\(\s*("|')ballerina\1\s*\)\s*\)/);
```



# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

--fcc-editable-region--

--fcc-editable-region--

const sellPlants = (plant, size, potsNo) => {
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}

console.log(sellPlants(ballerina, "small", 10));

console.log(catalog)
```
