---
id: 6734a606e5968a5a19b77171
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

前のステップと同様に、残りの植物をカタログに追加し続けてください。

# --hints--

`willowVale` オブジェクトを `catalog` マップのキーとして設定してください。

```js
assert.isTrue(catalog.has(willowVale));
```

`willowVale` マップの `catalog` キーの値は、`small`、`medium`、`large` のプロパティを持つオブジェクトで、それぞれの値は正の整数である必要があります。

```js
const b = catalog.get(willowVale)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

`hidcote` オブジェクトを `catalog` マップのキーとして設定してください。

```js
assert.isTrue(catalog.has(hidcote));
```

`hidcote` マップの `catalog` キーの値は、`small`、`medium`、`large` のプロパティを持つオブジェクトで、それぞれの値は正の整数である必要があります。

```js
const b = catalog.get(hidcote)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

`imperialGem` オブジェクトを `catalog` マップのキーとして設定してください。

```js
assert.isTrue(catalog.has(imperialGem));
```

`imperialGem` マップの `catalog` キーの値は、`small`、`medium`、`large` のプロパティを持つオブジェクトで、それぞれの値は正の整数である必要があります。

```js
const b = catalog.get(imperialGem)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

`royalCrown` オブジェクトを `catalog` マップのキーとして設定してください。

```js
assert.isTrue(catalog.has(royalCrown));
```

`royalCrown` マップの `catalog` キーの値は、`small`、`medium`、`large` のプロパティを持つオブジェクトで、それぞれの値は正の整数である必要があります。

```js
const b = catalog.get(royalCrown)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

`ballerina` マップにはまだ `catalog` キーが存在しているはずです。

```js
assert.isTrue(catalog.has(ballerina));
```

`prettyPolly` マップにはまだ `catalog` キーが存在しているはずです。

```js
assert.isTrue(catalog.has(prettyPolly));
```

`catalog` マップには6つの要素が含まれているはずです。

```js
assert.lengthOf(catalog, 6);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
