---
id: 6734a2385c52c54d916a8b86
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

カタログにもう1つキーと値のペアを追加してください。`prettyPolly` 変数をキーとして使い、値は再び整数の `small`、`medium`、`large` プロパティを持つオブジェクトにしてください。

# --hints--

`prettyPolly` オブジェクトを `catalog` マップのキーとして設定してください。

```js
assert.isTrue(catalog.has(prettyPolly));
```

`prettyPolly` マップの `catalog` キーの値は、`small`、`medium`、`large` プロパティを持つオブジェクトであるべきです。

```js
const b = catalog.get(prettyPolly)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

`small` マップの `prettyPolly` キーに関連付けられたオブジェクトの `catalog` プロパティは正の整数であるべきです。

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

`medium` マップの `prettyPolly` キーに関連付けられたオブジェクトの `catalog` プロパティは正の整数であるべきです。

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

`large` マップの `prettyPolly` キーに関連付けられたオブジェクトの `catalog` プロパティは正の整数であるべきです。

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

`ballerina` マップにはまだ `catalog` キーが存在しているはずです。

```js
assert.isTrue(catalog.has(ballerina));
```

`catalog` マップには2つの要素が含まれているはずです。

```js
assert.lengthOf(catalog, 2);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
