---
id: 67349ba230cff237c0f2bdf2
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

Mapオブジェクトは、キーがマップに挿入された順序を記憶します。マップオブジェクトにキーと値のペアを設定するには、`set`メソッドを使うことができます。以下は、`"name"`マップオブジェクトにキー`"Jerry"`を値`myMap`で設定しています。

```js
myMap.set("name", "Jerry");
```

`"ballerina"`マップの中で、`ballerina`キーを`catalog`オブジェクトの値に設定してください。その後、カタログをコンソールにログ出力してください。

マップオブジェクトを正しく見るには、ブラウザのコンソールを開く必要があることに注意してください。

# --hints--

`"ballerina"`マップの中で、`ballerina`キーを`catalog`オブジェクトの値に設定してください。

```js
assert.deepEqual(catalog.get("ballerina"), ballerina)
```

`catalog`マップには1つのキーと値のペアが含まれている必要があります。

```js
assert.lengthOf(catalog, 1);
```

`catalog`マップをコンソールにログ出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*catalog\s*\)/)
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();

--fcc-editable-region--
```
