---
id: 6734a13a29acdf47f2b8e161
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

通常のオブジェクトではキーは常に文字列でなければなりませんが、mapオブジェクトでは任意のデータ型をキーに使うことができます。カタログでは、植物オブジェクトをキーにして、異なるサイズの鉢の数を含むオブジェクトを値として関連付ける、オブジェクト同士のキーと値のペアを使います。

`set` の呼び出しを修正して、`ballerina` オブジェクトをキーとして使ってください。値には、`small`、`medium`、`large` の3つのプロパティを持つオブジェクトを使い、それぞれの値は正の整数にしてください。

# --hints--

`ballerina` オブジェクトを `catalog` マップのキーとして設定してください。

```js
assert.isTrue(catalog.has(ballerina));
```

`ballerina` マップの `catalog` キーの値は、`small`、`medium`、`large` のプロパティを持つオブジェクトであるべきです。

```js
const b = catalog.get(ballerina)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

`small` マップの `ballerina` キーに関連付けられたオブジェクトの `catalog` プロパティは正の整数であるべきです。

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

`medium` マップの `ballerina` キーに関連付けられたオブジェクトの `catalog` プロパティは正の整数であるべきです。

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

`large` マップの `ballerina` キーに関連付けられたオブジェクトの `catalog` プロパティは正の整数であるべきです。

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

`catalog` マップは1つのキーと値のペアを含むべきです。

```js
assert.lengthOf(catalog, 1);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set("ballerina", ballerina);
--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
