---
id: 69f0a37c3bf194a7c9c62492
title: Створити реєстр позик пристроїв
challengeType: 26
dashedName: lab-device-loan-ledger
---

# --description--

Ви — розробник програмного забезпечення, якому доручено створити застосунок-реєстр для керування ІТ-обладнанням, яке надає Сервісний центр вашої компанії. Реєстр — це об’єкт, ключами якого є теги ресурсів, а значеннями — об’єкти:

```js
const ledger = {
  "1": {
    type: "Laptop",
    status: "CheckedIn",
    borrower: {
      name: "",
      email: ""
    },
    dueDate: ""
  }
};
```

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історія користувача:**

1. Ви повинні змоделювати реєстр як об’єкт, де кожен пристрій ідентифікується своїм тегом ресурсу і має властивості `type`, `status`, `borrower` (з `name` та `email`) і `dueDate`.

1. Ви повинні реалізувати функцію `checkoutDevice(ledger, assetTag, borrower)`, яка приймає об’єкт реєстру і створює його копію без зміни оригіналу. У копії вона має встановити `name` та `email` об’єкта `borrower`, оновити `status` на `"CheckedOut"` і повернути об’єкт у форматі `{ ledger: updatedLedger, message: confirmationString }`, де `message` містить тег ресурсу та `name` об’єкта `borrower`.

   - Якщо `assetTag` не існує в реєстрі, функція має повернути незмінений реєстр разом із `message`, що містить тег ресурсу і пояснює, що його не знайдено.
   - Якщо пристрій уже має статус `"CheckedOut"`, функція має повернути незмінений реєстр, залишити поточного позичальника, а також повернути `message`, що містить тег ресурсу і пояснює, що пристрій уже позичений.

1. Ви повинні реалізувати функцію `checkinDevice(ledger, assetTag)`, яка приймає об’єкт реєстру і створює його копію без зміни оригіналу. У копії вона має очистити `name` та `email` об’єкта `borrower`, скинути `dueDate` до `""`, оновити `status` на `"CheckedIn"` і повернути об’єкт у форматі `{ ledger: updatedLedger, message: confirmationString }`, де `message` містить тег ресурсу.

   - Якщо `assetTag` не існує в реєстрі, функція має повернути незмінений реєстр разом із `message`, що містить тег ресурсу і пояснює, що його не знайдено.

1. Ви повинні реалізувати функцію `listOverdueDevices(ledger, today)`, яка приймає об’єкт ledger і повертає масив об’єктів пристроїв, що прострочені, відсортованих за `dueDate` у зростаючому порядку. Пристрій вважається простроченим, якщо його `status` дорівнює `"CheckedOut"`, а `dueDate` строго раніше за `today`. І `today`, і `dueDate` — це рядки у форматі місяць/день/рік, причому місяць і день можуть бути як з провідним нулем, так і без нього (наприклад, `"9/5/2025"` і `"09/05/2025"` обидва є дійсними).

1. Ви не повинні використовувати об’єкт JavaScript `Date` у своєму коді.

1. Ви повинні реалізувати функцію `serializeLedger(ledger)`, яка конвертує об’єкт реєстру у JSON-рядок, та функцію `loadLedger(json)`, яка конвертує JSON-рядок назад у JavaScript-об’єкт.

# --hints--

Ви повинні створити функцію з назвою `checkoutDevice`.

```js
assert.isFunction(checkoutDevice);
```

Ви повинні створити функцію з назвою `checkinDevice`.

```js
assert.isFunction(checkinDevice);
```

Ви повинні створити функцію з назвою `listOverdueDevices`.

```js
assert.isFunction(listOverdueDevices);
```

Ви повинні створити функцію з назвою `serializeLedger`.

```js
assert.isFunction(serializeLedger);
```

Ви повинні створити функцію з назвою `loadLedger`.

```js
assert.isFunction(loadLedger);
```

Ваша функція `checkoutDevice` не повинна змінювати оригінальний реєстр.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(ledger, original);
```

Ваша функція `checkoutDevice` має встановити `name` та `email` позичальника і встановити `status` у `"CheckedOut"`.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.isObject(result);
assert.equal(result.ledger["1"].status, "CheckedOut");
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
```

Ваша функція `checkoutDevice` має повертати підтверджувальне повідомлення, яке містить тег ресурсу та `name` позичальника.

```js
const ledger = {
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "2", { name: "Bob", email: "bob@test.com" });
assert.isString(result.message);
assert.include(result.message, "2");
assert.include(result.message, "Bob");
```

Ваша функція `checkoutDevice` має зберігати незмінними інші пристрої та поля у поверненому реєстрі.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

Ваша функція `checkoutDevice` має повертати незмінений реєстр і повідомлення, що містить тег ресурсу, якщо тег ресурсу не знайдено.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "99", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

Ваша функція `checkoutDevice` має повертати незмінений реєстр, залишати поточного позичальника і повертати повідомлення, що містить тег ресурсу, якщо пристрій уже позичений.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkoutDevice(ledger, "1", { name: "Bob", email: "bob@test.com" });
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "1");
```

Ваша функція `checkinDevice` не повинна змінювати оригінальний реєстр.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkinDevice(ledger, "1");
assert.deepEqual(ledger, original);
```

Ваша функція `checkinDevice` має очистити `name` та `email` позичальника, скинути `dueDate` до `""` і встановити `status` у `"CheckedIn"`.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].status, "CheckedIn");
assert.equal(result.ledger["1"].borrower.name, "");
assert.equal(result.ledger["1"].borrower.email, "");
assert.equal(result.ledger["1"].dueDate, "");
```

Ваша функція `checkinDevice` має повертати підтверджувальне повідомлення, яке містить тег ресурсу.

```js
const ledger = {
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Carol", email: "carol@test.com" }, dueDate: "11/15/2025" }
};
const result = checkinDevice(ledger, "3");
assert.isString(result.message);
assert.include(result.message, "3");
```

Ваша функція `checkinDevice` має зберігати незмінними інші пристрої та поля у поверненому реєстрі.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

Ваша функція `checkinDevice` має повертати незмінений реєстр і повідомлення, що містить тег ресурсу, якщо тег ресурсу не знайдено.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "99");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

Ваша функція `listOverdueDevices` має повертати масив.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
assert.isArray(listOverdueDevices(ledger, "12/01/2025"));
```

Ваша функція `listOverdueDevices` має повертати лише прострочені пристрої, у яких `status` дорівнює `"CheckedOut"`.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "12/31/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 1);
assert.equal(result[0].borrower.name, "Alice");
```

Ваша функція `listOverdueDevices` має повертати прострочені пристрої, відсортовані за `dueDate` у зростаючому порядку.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "11/01/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/01/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Bob");
assert.equal(result[1].borrower.name, "Alice");
```

Ваша функція `listOverdueDevices` має правильно сортувати навіть коли місяць або день не мають провідного нуля.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "9/5/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/1/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Alice");
assert.equal(result[1].borrower.name, "Bob");
```

Ви не повинні використовувати об’єкт JavaScript `Date` у своєму коді.

```js
assert.notMatch(__helpers.removeJSComments(code), /\bnew\s+Date\b|\bDate\s*[.(]/);
```

Ваша функція `serializeLedger` має повертати JSON-рядок, який представляє переданий реєстр.

```js
const obj = { "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" } };
assert.isString(serializeLedger(obj));
assert.deepEqual(JSON.parse(serializeLedger(obj)), obj);
```

Ваша функція `loadLedger` має повертати JavaScript-об’єкт, який представляє JSON-рядок.

```js
const json = '{"1":{"type":"Laptop","status":"CheckedIn","borrower":{"name":"","email":""},"dueDate":""}}';
assert.isObject(loadLedger(json));
assert.deepEqual(loadLedger(json), JSON.parse(json));
```

# --seed--

## --seed-contents--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

```

# --solutions--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

function serializeLedger(ledger) {
  return JSON.stringify(ledger);
}

function loadLedger(json) {
  return JSON.parse(json);
}

function checkoutDevice(ledger, assetTag, borrower) {
  const updatedLedger = JSON.parse(JSON.stringify(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  if (updatedLedger[assetTag].status === "CheckedOut") {
    return {
      ledger: updatedLedger,
      message: `${updatedLedger[assetTag].type} with asset tag ${assetTag} is already checked out.`
    };
  }

  updatedLedger[assetTag].status = "CheckedOut";
  updatedLedger[assetTag].borrower.name = borrower.name;
  updatedLedger[assetTag].borrower.email = borrower.email;
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked out to ${borrower.name}.`;
  return { ledger: updatedLedger, message };
}

function checkinDevice(ledger, assetTag) {
  const updatedLedger = JSON.parse(JSON.stringify(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  updatedLedger[assetTag].status = "CheckedIn";
  updatedLedger[assetTag].borrower.name = "";
  updatedLedger[assetTag].borrower.email = "";
  updatedLedger[assetTag].dueDate = "";
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked in.`;
  return { ledger: updatedLedger, message };
}

function reformatDate(date) {
  const datePartsArray = date.split("/");
  let month = datePartsArray[0];
  let day = datePartsArray[1];
  const year = datePartsArray[2];

  month = month.padStart(2, "0");
  day = day.padStart(2, "0");

  const formattedDate = year + month + day;
  return parseInt(formattedDate, 10);
}

function listOverdueDevices(ledger, today) {
  const todayFormatted = reformatDate(today);
  const overdue = [];

  for (const assetTag in ledger) {
    const item = ledger[assetTag];
    if (item.status === "CheckedOut" && item.dueDate) {
      if (reformatDate(item.dueDate) < todayFormatted) {
        overdue.push(item);
      }
    }
  }

  return overdue.sort((a, b) => reformatDate(a.dueDate) - reformatDate(b.dueDate));
}
```
