---
id: 669e81368e52b3a5c35a2dc5
title: 创建一个视频集页面
challengeType: 25
dashedName: build-a-video-compilation-page
demoType: onClick
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该有一个 `main` 元素，作为 `body` 元素的唯一子元素。
1. 你应该有一个包含页面主题的 `h1` 元素。
1. 你应该在 `h1` 元素下方有一个介绍页面主题的段落。
1. 你应该在第一个段落下方添加三个 `section` 元素。
1. 每个部分应该依次包含一个 `h2` 元素、一个段落和一个 `iframe` 元素。
1. 三个 `iframe` 都应该有设置为有效视频的 `src` 属性。
1. 每个 `iframe` 元素也都应该有一个用于描述嵌入内容的 `title` 属性，以及设置为合适大小的 `height` 属性和 `width` 属性。

# --hints--

你应该在 `body` 元素中包含一个 `main` 元素。

```js
assert.equal(document.querySelector('body :first-child').tagName, 'MAIN');
```

你的 `main` 元素应该是 `body` 元素的唯一子元素。

```js
assert.equal(document.querySelector('body').children.length, 1);
```

你应该在 `main` 元素中加入一个 `h1` 元素，用于页面主题。

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

你应该在 `h1` 元素下方有一个介绍页面主题的段落。

```js
const p = document.querySelectorAll('p')[0];
assert.equal(p?.previousElementSibling?.tagName, 'H1');
assert.isAbove(p?.innerText.length, 0);
```

你应该在第一个 `p` 元素下方添加三个 `section` 元素。

```js
assert.equal(document.querySelectorAll('section')?.length, 3);
assert.equal(document.querySelectorAll('section')[0]?.previousElementSibling?.tagName, 'P')
```

每个 `section` 元素都应该以一个 `h2` 开头，作为该部分的标题。

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[0]?.tagName, 'H2');
    assert.isAbove(section.children[0]?.innerText.length, 0);
}
```

每个 `section` 元素都应该包含一个 `p` 元素作为它的第二个子元素，用于介绍视频内容。

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[1]?.tagName, 'P');
    assert.isAbove(section.children[1]?.innerText.length, 0);
}
```

每个 `section` 元素都应该包含一个 `iframe` 元素，作为它的第三个子元素。

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[2]?.tagName, 'IFRAME');
}
```

你的第一个 `iframe` 元素应该有一个设置为有效视频的 `src` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.src.length, 0)
```

你的第一个 `iframe` 元素应该有一个描述嵌入内容的 `title` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.title.length, 0)
```

你的第一个 `iframe` 元素应该有一个 `height` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.height.length, 0)
```

你的第一个 `iframe` 元素应该有一个 `width` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.width.length, 0)
```

你的第二个 `iframe` 元素应该有一个设置为有效视频的 `src` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.src.length, 0)
```

你的第二个 `iframe` 元素应该有一个描述嵌入内容的 `title` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.title.length, 0)
```

你的第二个 `iframe` 元素应该有一个 `height` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.height.length, 0)
```

你的第二个 `iframe` 元素应该有一个 `width` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.width.length, 0)
```

你的第三个 `iframe` 元素应该有一个设置为有效视频的 `src` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.src.length, 0)
```

你的第三个 `iframe` 元素应该有一个描述嵌入内容的 `title` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.title.length, 0)
```

你的第三个 `iframe` 元素应该有一个 `height` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.height.length, 0)
```

你的第三个 `iframe` 元素应该有一个 `width` 属性。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.width.length, 0)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>
    <main>
        <h1>Front End Web Development</h1>
        <p>Front End Web Development involves designing and building user interfaces of websites using HTML, CSS, and
            JavaScript. It focuses on creating visually appealing, interactive, and user-friendly websites. Front End
            developers
            ensure accessibility, performance optimization, and
            adherence to web standards, blending creativity with technical expertise.</p>
        <section>
            <h2>HTML</h2>
            <p>HTML, or HyperText Markup Language, forms the backbone of web pages by defining their basic structure and
                layout. It uses a series of elements to display text, images, lists, tables, and other static content on
                a web page.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/GDGejH3SDNQ?si=KJYLgcz4kyyroYMB"
                title="html and coding introduction video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>CSS</h2>
            <p>CSS, or Cascading Style Sheets, is used to control the visual presentation of HTML elements on a web
                page. It allows for customized styling, including colors, fonts, layouts, and spacing, making websites
                visually appealing. CSS also supports responsive design, ensuring web pages look good on
                all devices.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/OXGznpKZ_sA?si=s3KDSZvrhU_PU9XL"
                title="css tutorial video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>JavaScript</h2>
            <p>JavaScript is a programming language that adds interactivity to static web pages, enabling dynamic
                content updates, form validations, animations, and responsive behaviors based on user interactions and
                events.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/jS4aFq5-91M?si=zKQhHEYwU4tnMmVm"
                title="javascript programming video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
    </main>
</body>

</html>
```
