---
id: 6852db755c9c1274d75aae9a
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому практичному занятті ви створите витягач PIN-коду. Цифри PIN приховані в кожній лінії вірша.

Почніть зі створення функції з назвою `pin_extractor` з параметром `poem`. Функція без тіла викликає помилку, тому додайте `pass` всередину функції як тимчасовий заповнювач, щоб зробити її дійсною.

# --hints--

У вас має бути функція з назвою `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('pin_extractor')
    `)
})
```

Функція `pin_extractor` має мати параметр з назвою `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poem')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
