---
id: 6852dd8486c7ee0171322950
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Вам потрібна змінна, щоб зберігати секретний код під час його декодування, тому видаліть `pass` і створіть змінну з назвою `secret_code` всередині функції, призначивши їй порожній рядок.

# --hints--

У вас має бути змінна `secret_code` всередині функції `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_code')
    `)
})
```

Змінна `secret_code` має мати значення порожнього рядка.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_code').is_equivalent('secret_code = ""')
    `)
})
```

У вашому коді більше не повинно бути `pass`.

```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').has_pass()
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
