---
id: 68530231c05b93324e071104
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Поки ви працюєте над функцією, було б корисно мати виклик функції для налагодження, тому створіть глобальну змінну `poem` поза межами функції, яка містить цей вірш:

```md
Stars and the moon
shine in the sky
white and bright
until the end of the night
```

Для цього можна використати багаторядковий рядок.

Потім викличте функцію з аргументом `poem`.

# --hints--

У вас має бути глобальна змінна з іменем `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_variable('poem')
    `)
})
```

Значення змінної `poem` має дорівнювати `Stars and the moon\nshine in the sky\nwhite and bright\nuntil the end of the night`.

```js
({
    test: () => runPython(`
    pm = """Stars and the moon
    shine in the sky
    white and bright
    until the end of the night"""
    assert _Node(_code).find_variable('poem').is_equivalent(f'poem = """{pm}"""')
    `)
})
```

Ви маєте викликати функцію `pin_extractor` з аргументом `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('pin_extractor(poem)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''

--fcc-editable-region--

--fcc-editable-region--
```
