---
id: 68531fbb934d735415053274
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Тепер, коли ви додали виклик функції, ви можете бачити будь-яку помилку, створену всередині функції, у терміналі, і завжди можете використати `print` щоб побачити значення речей.

У кожній лінії прихована цифра пін-коду, тому всередині функції використайте метод `split`, щоб розбити рядок на список ліній. Розбийте лінії за символом нового рядка (`\n`) і призначте отриманий список змінній з ім’ям `lines`.

# --hints--

У вас має бути змінна з ім’ям `lines` у вашій функції `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('lines')
    `)
})
```

Змінна `lines` має мати значення `poem.split('\n')`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('lines').is_equivalent('lines = poem.split("\\\\n")')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
```
