---
id: 685578aae65e4106ad2b4300
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Як ви дізналися в попередньому уроці, функція `enumerate` дозволяє відстежувати індекс під час ітерування над ітерованим об’єктом:

```py
fruits = ['apple', 'plum', 'bananas']

for index, item in enumerate(fruits):
    print(index, item)

# 0 apple
# 1 plum
# 2 bananas
```

n-та цифра секретного коду походить з n-го слова n-ої лінії, тому потрібно знати номер кожної лінії.

Змініть цикл так, щоб він ітерував над `enumerate(lines)`, і додайте ще одну змінну циклу перед `line` з ім’ям `line_index`.

Також оновіть виклик `print` на `print(line_index, line)`.

# --hints--

Цикл має ітерувати над `enumerate(lines)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('enumerate(lines)')
    `)
})
```

Змінна циклу має бути `line_index, line`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line_index, line')
    `)
})
```

У циклі має бути `print(line_index, line)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line_index, line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    for line in lines:
        print(line)
--fcc-editable-region--
        words = line.split()
        print(words)

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
