---
id: 68559f3561d7df287586142e
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

Тепер екстрактор пін-кодів працює, але він міг би бути ефективнішим! Він міг би одночасно витягувати пін-коди з багатьох віршів!

Додайте ще два вірші в глобальну область видимості, `poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'` і `poem3 = 'There\nonce\nwas\na\ndragon'`, а поки що перетворіть на коментарі `print` з викликом `pin_extractor`.


# --hints--

Змінні `poem2` і `poem3` мають існувати і мати правильне значення.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('poem2')
        assert _Node(_code).has_variable('poem3')
        assert _Node(_code).find_variable('poem2').is_equivalent("poem2 = 'The grass is green\\\\nhere and there\\\\nhoping for rain\\\\nbefore it turns yellow'")
        assert _Node(_code).find_variable('poem3').is_equivalent("poem3 = 'There\\\\nonce\\\\nwas\\\\na\\\\ndragon'")
    `)
})
```

Ви маєте перетворити на коментарі виклик `print(pin_extractor(poem))`.

```js
({
    test: () => runPython(`
        assert not _Node(_code).has_call('print(pin_extractor(poem))')`)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
        else:
            secret_code += '0'
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

--fcc-editable-region--


print(pin_extractor(poem))
--fcc-editable-region--

```
