---
id: 6855918e87c33c6979d21d3f
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Зараз функція добре справляється з вилученням коду з вірша, але є крайній випадок, який потрібно врахувати: рядок вірша може бути коротшим, ніж очікувалося, і це спричинить помилку у функції (ви можете спробувати це, видаливши слово `bright` з вірша).

Помістіть лінію `secret_code += str(len(words[line_index]))` в оператор `if`, який перевіряє, що в списку `words` достатньо слів.

# --hints--

У циклі має бути оператор `if`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0]
    `)
})
```

Умова оператора if має бути `len(words) > line_index`.


```js
({
    test: () => runPython(`
    _cond = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
    assert _cond.is_equivalent('len(words) > line_index') or _cond.is_equivalent('line_index < len(words)')
    `)
})
```

Всередині оператора `if` має бути `secret_code += str(len(words[line_index]))`.


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt('secret_code += str(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
--fcc-editable-region--
        secret_code += str(len(words[line_index]))
--fcc-editable-region--
    return secret_code

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

print(pin_extractor(poem))

```
