---
id: 685535bf527b2c77236b3a9a
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

n-та цифра пін-коду прихована як довжина n-го слова в n-тій лінії.

Щоб знайти довжину n-го слова, наступним кроком є розбити лінію вірша на список слів: всередині циклу створіть змінну `words` і призначте їй значення `line`, розбитого на слова за допомогою методу `split`.

Потім додайте виклик `print` з `words` як аргументом.

# --hints--

У вас має бути змінна `words` всередині циклу.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_variable('words')
    `)
})
```

Значення змінної `words` має бути `line.split()`.

```js
({
    test: () => runPython(`
    variable = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_variable('words')
    assert variable.is_equivalent('words = line.split()') or variable.is_equivalent('words = line.split(" ")')
    `)
})
```

У вас має бути `print(words)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        print(line)
--fcc-editable-region--
        
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
