---
id: 68557bd29f7f3c1880f98836
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

В останній лінії циклу оновіть виклик `print`, щоб він виводив слово всередині списку `words` за індексом `line_index`.

# --hints--

У вас має бути `print(words[line_index])` наприкінці циклу.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words[line_index])')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
--fcc-editable-region--
        print(words)
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
