---
id: 68553423d775cc60134ac8bb
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Вам потрібно працювати з кожною лінією окремо: створіть цикл над `lines`, який використовує `line` як змінну циклу. Усередині циклу виведіть `line`.

# --hints--

У вашій функції `for` має бути цикл `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0]
    `)
})
```

Цикл має ітерувати над `lines`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('lines')
    `)
})
```

Змінною циклу має бути `line` (`for line in lines`).

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line')
    `)
})
```

Усередині циклу має бути `print(line)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
