---
id: 587d7da9367417b2b2512b66
title: Unganisha Safu Mbili za Data Kutumia Njia ya concat
challengeType: 1
forumTopicId: 301229
dashedName: combine-two-arrays-using-the-concat-method
---

# --description--

<dfn>Concatenation</dfn> inamaanisha kuunganisha vitu mwisho kwa mwisho. JavaScript hutoa njia ya `concat` kwa mfuatano wa herufi na safu za data ambayo hufanya kazi kwa njia ile ile. Kwa safu za data, njia hii inaitwa kwenye moja, kisha safu nyingine hutolewa kama hoja kwa `concat`, ambayo inaongezwa mwishoni mwa safu ya kwanza. Inarudisha safu mpya na haisababishi mabadiliko yoyote kwa safu za awali. Hapa kuna mfano:

```js
[1, 2, 3].concat([4, 5, 6]);
```

Safu iliyorejeshwa itakuwa `[1, 2, 3, 4, 5, 6]`.

# --instructions--

Tumia njia ya `concat` katika kitendakazi cha `nonMutatingConcat` kuunganisha `attach` mwishoni mwa `original`. Kitendakazi kinapaswa kurudisha safu iliyounganishwa.

# --hints--

Msimbo wako unapaswa kutumia njia ya `concat`.

```js
assert(__helpers.removeJSComments(code).match(/\.concat/g));
```

Safu ya `first` haipaswi kubadilika.

```js
assert(JSON.stringify(first) === JSON.stringify([1, 2, 3]));
```

Safu ya `second` haipaswi kubadilika.

```js
assert(JSON.stringify(second) === JSON.stringify([4, 5]));
```

`nonMutatingConcat([1, 2, 3], [4, 5])` inapaswa kurudisha `[1, 2, 3, 4, 5]`.

```js
assert(
  JSON.stringify(nonMutatingConcat([1, 2, 3], [4, 5])) ===
    JSON.stringify([1, 2, 3, 4, 5])
);
```

# --seed--

## --seed-contents--

```js
function nonMutatingConcat(original, attach) {
  // Only change code below this line


  // Only change code above this line
}

const first = [1, 2, 3];
const second = [4, 5];
nonMutatingConcat(first, second);
```

# --solutions--

```js
function nonMutatingConcat(original, attach) {
  return original.concat(attach);
}
const first = [1, 2, 3];
const second = [4, 5];
```
