---
id: 587d7b8d367417b2b2512b5b
title: Jifunze kuhusu programu utendaji
challengeType: 1
forumTopicId: 301233
dashedName: learn-about-functional-programming
---

# --description--

Programu utendaji ni mtindo wa programu ambapo suluhisho ni vitendakazi rahisi, vilivyo huru, bila athari za pembeni nje ya wigo wa kitendakazi: `INPUT -> PROCESS -> OUTPUT`

Programu utendaji ni kuhusu:

1) Vitendakazi vilivyo huru - hakuna utegemezi kwa hali ya programu, ambayo ni pamoja na vigezo vya kimataifa vinavyoweza kubadilika

2) Vitendakazi safi - ingizo lile lile daima hutoa matokeo yale yale

3) Vitendakazi vyenye athari za pembeni chache - mabadiliko yoyote, au mabadiliko, kwa hali ya programu nje ya kitendakazi yanadhibitiwa kwa makini

# --instructions--

Wanachama wa freeCodeCamp wanapenda chai.

Katika mhariri wa msimbo, vitendakazi vya `prepareTea` na `getTea` vimeshaainishwa kwako. Fanya wito wa kitendakazi cha `getTea` kupata vikombe 40 vya chai kwa timu, na vihifadhi katika kigezo cha `tea4TeamFCC`.

# --hints--

Kigezo cha `tea4TeamFCC` kinapaswa kuhifadhi vikombe 40 vya chai kwa timu.

```js
assert(tea4TeamFCC.length === 40);
```

Kigezo cha `tea4TeamFCC` kinapaswa kuhifadhi vikombe vya chai ya kijani.

```js
assert(tea4TeamFCC[0] === 'greenTea');
```

# --seed--

## --seed-contents--

```js
// Function that returns a string representing a cup of green tea
const prepareTea = () => 'greenTea';

/*
Given a function (representing the tea type) and number of cups needed, the
following function returns an array of strings (each representing a cup of
a specific type of tea).
*/
const getTea = (numOfCups) => {
  const teaCups = [];

  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }
  return teaCups;
};

// Only change code below this line
const tea4TeamFCC = null;
// Only change code above this line
```

# --solutions--

```js
const prepareTea = () => 'greenTea';

const getTea = (numOfCups) => {
  const teaCups = [];
  
  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }

  return teaCups;
};

const tea4TeamFCC = getTea(40); 
```
