---
id: 587d7b90367417b2b2512b65
title: Rudisha Sehemu ya Safu ya Data Kutumia Njia ya slice
challengeType: 1
forumTopicId: 301239
dashedName: return-part-of-an-array-using-the-slice-method
---

# --description--

Njia ya `slice` inarudisha nakala ya vipengele fulani vya safu ya data. Inaweza kuchukua hoja mbili, ya kwanza hutoa kielezo cha mahali pa kuanza kukata, ya pili ni kielezo cha mahali pa kumaliza kukata (na haijumuishi). Ikiwa hoja hazitolewi, chaguo-msingi ni kuanza mwanzoni mwa safu ya data hadi mwisho, ambayo ni njia rahisi ya kutengeneza nakala ya safu nzima. Njia ya `slice` haibadilishi safu ya data ya awali, bali inarudisha mpya.

Hapa kuna mfano:

```js
const arr = ["Cat", "Dog", "Tiger", "Zebra"];
const newArray = arr.slice(1, 3);
```

`newArray` ingekuwa na thamani ya `["Dog", "Tiger"]`.

# --instructions--

Tumia njia ya `slice` katika kitendakazi cha `sliceArray` kurudisha sehemu ya safu ya data ya `anim` ikizingatiwa viashiria vya `beginSlice` na `endSlice` vilivyotolewa. Kitendakazi kinapaswa kurudisha safu ya data.

# --hints--

Msimbo wako unapaswa kutumia njia ya `slice`.

```js
assert(__helpers.removeJSComments(code).match(/\.slice/g));
```

Kigezo cha `inputAnim` hakipaswi kubadilika.

```js
assert(
  JSON.stringify(inputAnim) ===
    JSON.stringify(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 1, 3)` inapaswa kurudisha `["Dog", "Tiger"]`.

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 1, 3)) ===
    JSON.stringify(['Dog', 'Tiger'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 0, 1)` inapaswa kurudisha `["Cat"]`.

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 0, 1)) ===
    JSON.stringify(['Cat'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 1, 4)` inapaswa kurudisha `["Dog", "Tiger", "Zebra"]`.

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 1, 4)) ===
    JSON.stringify(['Dog', 'Tiger', 'Zebra'])
);
```

# --seed--

## --seed-contents--

```js
function sliceArray(anim, beginSlice, endSlice) {
  // Only change code below this line


  // Only change code above this line
}

const inputAnim = ["Cat", "Dog", "Tiger", "Zebra", "Ant"];
sliceArray(inputAnim, 1, 3);
```

# --solutions--

```js
function sliceArray(anim, beginSlice, endSlice) {
  return anim.slice(beginSlice, endSlice);
}
const inputAnim = ["Cat", "Dog", "Tiger", "Zebra", "Ant"];
```
