---
id: 587d7da9367417b2b2512b69
title: Panga safu ya data kwa mpangilio wa alfabeti kwa kutumia njia ya sort
challengeType: 1
forumTopicId: 18303
dashedName: sort-an-array-alphabetically-using-the-sort-method
---

# --description--

Njia ya `sort` hupanga vipengele vya safu ya data kulingana na kitendakazi cha mwitiko.

Kwa mfano:

```js
function ascendingOrder(arr) {
  return arr.sort(function(a, b) {
    return a - b;
  });
}

ascendingOrder([1, 5, 2, 3, 4]);
```

Hii itarudisha thamani ya `[1, 2, 3, 4, 5]`.

```js
function reverseAlpha(arr) {
  return arr.sort(function(a, b) {
    return a === b ? 0 : a < b ? 1 : -1;
  });
}

reverseAlpha(['l', 'h', 'z', 'b', 's']);
```

Hii itarudisha thamani ya `['z', 's', 'l', 'h', 'b']`.

Njia ya kawaida ya kupanga katika JavaScript ni kwa thamani ya pointi ya Unicode ya mfuatano wa herufi, ambayo inaweza kurudisha matokeo yasiyotegemewa. Kwa hiyo, inahimizwa kutoa kitendakazi cha mwitiko ili kubainisha jinsi ya kupanga vipengele vya safu ya data. Kitendakazi hicho cha mwitiko, kinachoitwa kawaida `compareFunction`, kinapowekwa, vipengele vya safu ya data hupangwa kulingana na thamani ya return ya `compareFunction`: Ikiwa `compareFunction(a,b)` inarudisha thamani ndogo kuliko 0 kwa vipengele viwili `a` na `b`, basi `a` itakuja kabla ya `b`. Ikiwa `compareFunction(a,b)` inarudisha thamani kubwa kuliko 0 kwa vipengele viwili `a` na `b`, basi `b` itakuja kabla ya `a`. Ikiwa `compareFunction(a,b)` inarudisha thamani sawa na 0 kwa vipengele viwili `a` na `b`, basi `a` na `b` zitabaki kama zilivyo.

# --instructions--

Tumia njia ya `sort` katika kitendakazi cha `alphabeticalOrder` kupanga vipengele vya `arr` kwa mpangilio wa alfabeti. Kitendakazi kinapaswa kurudisha safu ya data iliyopangwa.

# --hints--

Msimbo wako unapaswa kutumia njia ya `sort`.

```js
assert(__helpers.removeJSComments(code).match(/\.sort/g));
```

`alphabeticalOrder(["a", "d", "c", "a", "z", "g"])` inapaswa kurudisha `["a", "a", "c", "d", "g", "z"]`.

```js
assert(
  JSON.stringify(alphabeticalOrder(['a', 'd', 'c', 'a', 'z', 'g'])) ===
    JSON.stringify(['a', 'a', 'c', 'd', 'g', 'z'])
);
```

`alphabeticalOrder(["x", "h", "a", "m", "n", "m"])` inapaswa kurudisha `["a", "h", "m", "m", "n", "x"]`.

```js
assert(
  JSON.stringify(alphabeticalOrder(['x', 'h', 'a', 'm', 'n', 'm'])) ===
    JSON.stringify(['a', 'h', 'm', 'm', 'n', 'x'])
);
```

`alphabeticalOrder(["a", "a", "a", "a", "x", "t"])` inapaswa kurudisha `["a", "a", "a", "a", "t", "x"]`.

```js
assert(
  JSON.stringify(alphabeticalOrder(['a', 'a', 'a', 'a', 'x', 't'])) ===
    JSON.stringify(['a', 'a', 'a', 'a', 't', 'x'])
);
```

# --seed--

## --seed-contents--

```js
function alphabeticalOrder(arr) {
  // Only change code below this line

  return arr
  // Only change code above this line
}

alphabeticalOrder(["a", "d", "c", "a", "z", "g"]);
```

# --solutions--

```js
function alphabeticalOrder(arr) {
  return arr.sort();
}
```
