---
id: 587d7b8f367417b2b2512b62
title: Tekeleza map kwenye mfano wa prototype
challengeType: 1
forumTopicId: 301230
dashedName: implement-map-on-a-prototype
---

# --description--

Kama ulivyoona kutoka kwa kutumia `Array.prototype.map()`, au kwa urahisi `map()` hapo awali, njia ya `map` inarudisha safu ya data yenye urefu sawa na ile iliyoitwa. Pia haibadilishi safu ya data ya awali, mradi tu kitendakazi chake cha mwitiko hakibadilishi.

Kwa maneno mengine, `map` ni kitendakazi safi, na matokeo yake hutegemea kabisa ingizo lake. Zaidi ya hayo, huchukua kitendakazi kingine kama hoja yake.

Unaweza kujifunza mengi kuhusu njia ya `map` ukitekeleza toleo lako mwenyewe la njia hiyo. Inashauriwa utumie mzunguko wa `for` au `Array.prototype.forEach()`.

# --instructions--

Andika `Array.prototype.myMap()` yako mwenyewe, ambayo inapaswa kufanya kazi sawa kabisa na `Array.prototype.map()`. Hupaswi kutumia njia ya `map` iliyojengwa ndani. Mfano wa `Array` unaweza kufikiwa katika njia ya `myMap` kwa kutumia `this`.

# --hints--

`[23, 65, 98, 5, 13].myMap(item => item * 2)` inapaswa kuwa sawa na `[46, 130, 196, 10, 26]`.

```js
const _test_s = [23, 65, 98, 5, 13];
const _callback = item => item * 2;
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

`["naomi", "quincy", "camperbot"].myMap(element => element.toUpperCase())` inapaswa kurudisha `["NAOMI", "QUINCY", "CAMPERBOT"]`.

```js
const _test_s = ["naomi", "quincy", "camperbot"];
const _callback = element => element.toUpperCase();
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

`[1, 1, 2, 5, 2].myMap((element, index, array) => array[index + 1] || array[0])` inapaswa kurudisha `[1, 2, 5, 2, 1]`.

```js
const _test_s = [1, 1, 2, 5, 2];
const _callback = (element, index, array) => array[index + 1] || array[0];
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

Msimbo wako haupaswi kutumia njia ya `map`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?map/g));
```

# --seed--

## --seed-contents--

```js
Array.prototype.myMap = function(callback) {
  const newArray = [];
  // Only change code below this line

  // Only change code above this line
  return newArray;
};
```

# --solutions--

```js
Array.prototype.myMap = function(callback) {
  const newArray = [];
  for (let i = 0; i < this.length; i++) {
    newArray.push(callback(this[i], i, this));
  }
  return newArray;
};

// Test case
const s = [23, 65, 98, 5];
const doubled_s = s.myMap(item => item * 2);
```
