---
id: 587d7dab367417b2b2512b6f
title: Tumia njia ya some kuangalia kama vipengele vyovyote katika safu ya data vinakidhi vigezo
challengeType: 1
forumTopicId: 301314
dashedName: use-the-some-method-to-check-that-any-elements-in-an-array-meet-a-criteria
---

# --description--

Njia ya `some` hufanya kazi na safu za data kuangalia kama *kivipi* kipengele chochote kinapita mtihani fulani. Inarudisha thamani ya kweli au siyo kweli - `true` ikiwa thamani yoyote inakidhi vigezo, `false` ikiwa sivyo.

Kwa mfano, msimbo ufuatao ungeangalia kama kipengele chochote katika safu ya data ya `numbers` ni kidogo kuliko 10:

```js
const numbers = [10, 50, 8, 220, 110, 11];

numbers.some(function(currentValue) {
  return currentValue < 10;
});
```

Njia ya `some` ingetumia `true`.

# --instructions--

Tumia njia ya `some` ndani ya kitendakazi cha `checkPositive` kuangalia kama kipengele chochote katika `arr` ni chanya. Kitendakazi kinapaswa kurudisha thamani ya kweli au siyo kweli.

# --hints--

Msimbo wako unapaswa kutumia njia ya `some`.

```js
assert(__helpers.removeJSComments(code).match(/\.some/g));
```

`checkPositive([1, 2, 3, -4, 5])` inapaswa kurudisha `true`.

```js
assert(checkPositive([1, 2, 3, -4, 5]));
```

`checkPositive([1, 2, 3, 4, 5])` inapaswa kurudisha `true`.

```js
assert(checkPositive([1, 2, 3, 4, 5]));
```

`checkPositive([-1, -2, -3, -4, -5])` inapaswa kurudisha `false`.

```js
assert(!checkPositive([-1, -2, -3, -4, -5]));
```

# --seed--

## --seed-contents--

```js
function checkPositive(arr) {
  // Only change code below this line


  // Only change code above this line
}

checkPositive([1, 2, 3, -4, 5]);
```

# --solutions--

```js
function checkPositive(arr) {
  return arr.some(elem => elem > 0);
}
```
