---
id: 587d7daa367417b2b2512b6c
title: Unganisha safu ya data kuwa mfuatano wa herufi kwa kutumia njia ya join
challengeType: 1
forumTopicId: 18221
dashedName: combine-an-array-into-a-string-using-the-join-method
---

# --description--

Njia ya `join` hutumika kuunganisha vipengele vya safu ya data pamoja ili kuunda mfuatano wa herufi. Inachukua hoja ya kitenzi cha kutenganisha kinachotumika kutenganisha vipengele vya safu ya data ndani ya mfuatano wa herufi.

Hapa kuna mfano:

```js
const arr = ["Hello", "World"];
const str = arr.join(" ");
```

`str` ingekuwa na thamani ya mfuatano wa herufi `Hello World`.
# --instructions--

Tumia njia ya `join` (miongoni mwa nyingine) ndani ya kitendakazi cha `sentensify` kutengeneza sentensi kutoka kwa maneno yaliyomo katika mfuatano wa herufi `str`. Kitendakazi kinapaswa kurudisha mfuatano wa herufi. Kwa mfano, `I-like-Star-Wars` itageuzwa kuwa `I like Star Wars`. Kwa changamoto hii, usitumie njia ya `replace`.

# --hints--

Msimbo wako unapaswa kutumia njia ya `join`.

```js
assert(__helpers.removeJSComments(code).match(/\.join/g));
```

Msimbo wako haupaswi kutumia njia ya `replace`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?replace/g));
```

`sentensify("May-the-force-be-with-you")` inapaswa kurudisha mfuatano wa herufi.

```js
assert(typeof sentensify('May-the-force-be-with-you') === 'string');
```

`sentensify("May-the-force-be-with-you")` inapaswa kurudisha mfuatano wa herufi `May the force be with you`.

```js
assert(sentensify('May-the-force-be-with-you') === 'May the force be with you');
```

`sentensify("The.force.is.strong.with.this.one")` inapaswa kurudisha mfuatano wa herufi `The force is strong with this one`.

```js
assert(
  sentensify('The.force.is.strong.with.this.one') ===
    'The force is strong with this one'
);
```

`sentensify("There,has,been,an,awakening")` inapaswa kurudisha mfuatano wa herufi `There has been an awakening`.

```js
assert(
  sentensify('There,has,been,an,awakening') === 'There has been an awakening'
);
```

# --seed--

## --seed-contents--

```js
function sentensify(str) {
  // Only change code below this line


  // Only change code above this line
}

sentensify("May-the-force-be-with-you");
```

# --solutions--

```js
function sentensify(str) {
  return str.split(/\W/).join(' ');
}
```
