---
id: 587d7daa367417b2b2512b6b
title: Gawanya mfuatano wa herufi kuwa safu ya data kwa kutumia njia ya split
challengeType: 1
forumTopicId: 18305
dashedName: split-a-string-into-an-array-using-the-split-method
---

# --description--

Njia ya `split` hugawanya mfuatano wa herufi kuwa safu ya data za mfuatano wa herufi. Inachukua hoja ya kigezo cha mgawanyiko, ambayo inaweza kuwa herufi inayotumika kugawanya mfuatano wa herufi au usemi wa kawaida. Kwa mfano, ikiwa mgawanyiko ni nafasi, unapata safu ya maneno, na ikiwa mgawanyiko ni mfuatano wa herufi tupu, unapata safu ya kila herufi katika mfuatano wa herufi.

Hapa kuna mifano miwili inayogawanya mfuatano mmoja kwa nafasi, kisha mwingine kwa nambari kwa kutumia usemi wa kawaida:

```js
const str = "Hello World";
const bySpace = str.split(" ");

const otherString = "How9are7you2today";
const byDigits = otherString.split(/\d/);
```

`bySpace` itakuwa na thamani `["Hello", "World"]` na `byDigits` itakuwa na thamani `["How", "are", "you", "today"]`.

Kwa kuwa mfuatano wa herufi hauwezi kubadilika, njia ya `split` inafanya iwe rahisi kufanya kazi nao.

# --instructions--

Tumia njia ya `split` ndani ya kitendakazi cha `splitify` kugawanya `str` kuwa safu ya maneno. Kitendakazi kinapaswa kurudisha safu hiyo. Kumbuka kwamba maneno hayatenganishwi kila mara kwa nafasi, na safu haipaswi kuwa na alama za uandishi.

# --hints--

Msimbo wako unapaswa kutumia njia ya `split`.

```js
assert(__helpers.removeJSComments(code).match(/\.split/g));
```

`splitify("Hello World,I-am code")` inapaswa kurudisha `["Hello", "World", "I", "am", "code"]`.

```js
assert(
  JSON.stringify(splitify('Hello World,I-am code')) ===
    JSON.stringify(['Hello', 'World', 'I', 'am', 'code'])
);
```

`splitify("Earth-is-our home")` inapaswa kurudisha `["Earth", "is", "our", "home"]`.

```js
assert(
  JSON.stringify(splitify('Earth-is-our home')) ===
    JSON.stringify(['Earth', 'is', 'our', 'home'])
);
```

`splitify("This.is.a-sentence")` inapaswa kurudisha `["This", "is", "a", "sentence"]`.

```js
assert(
  JSON.stringify(splitify('This.is.a-sentence')) ===
    JSON.stringify(['This', 'is', 'a', 'sentence'])
);
```

# --seed--

## --seed-contents--

```js
function splitify(str) {
  // Only change code below this line


  // Only change code above this line
}

splitify("Hello World,I-am code");
```

# --solutions--

```js
function splitify(str) {
  return str.split(/\W/);
}
```
