---
id: 587d7b88367417b2b2512b45
title: 'Tumia kitendakazi ya daraja la juu map, filter, au reduce kutatua changamoto ngumu'
challengeType: 1
forumTopicId: 301311
dashedName: use-higher-order-functions-map-filter-or-reduce-to-solve-a-complex-problem
---

# --description--

Sasa baada ya kufanya changamoto chache ukitumia vitendakazi vya daraja la juu kama `map()`, `filter()`, na `reduce()`, sasa unaweza kuvitumia kutatua changamoto ngumu zaidi.

# --instructions--

Kamilisha msimbo wa kitendakazi cha `squareList` ukitumia mchanganyiko wowote wa `map()`, `filter()`, na `reduce()`. Kitendakazi kinapaswa kurudisha safu mpya ya data yenye mraba wa *nambari kamili chanya tu* (nambari za desimali si nambari kamili) wakati safu ya nambari halisi inapokelewa. Mfano wa safu ya nambari halisi ni `[-3, 4.8, 5, 3, -3.2]`.

**Kumbuka:** Kitendakazi chako hakipaswi kutumia aina yoyote ya `for` au `while` au kitendakazi cha `forEach()`.

# --hints--

`squareList` inapaswa kuwa `function`.

```js
assert.typeOf(squareList, 'function'),
  '<code>squareList</code> should be a <code>function</code>';
```

`for`, `while`, na `forEach` hazipaswi kutumika.

```js
assert(!__helpers.removeJSComments(code).match(/for|while|forEach/g));
```

`map`, `filter`, au `reduce` zinapaswa kutumika.

```js
assert(
  __helpers
    .removeWhiteSpace(__helpers.removeJSComments(code))
    .match(/\.(map|filter|reduce)\(/g)
);
```

Kitendakazi kinapaswa kurudisha `array`.

```js
assert(Array.isArray(squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2])));
```

`squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2])` inapaswa kurudisha `[16, 1764, 36]`.

```js
assert.deepStrictEqual(squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2]), [
  16,
  1764,
  36
]);
```

`squareList([-3.7, -5, 3, 10, 12.5, 7, -4.5, -17, 0.3])` inapaswa kurudisha `[9, 100, 49]`.

```js
assert.deepStrictEqual(squareList([-3.7, -5, 3, 10, 12.5, 7, -4.5, -17, 0.3]), [
  9,
  100,
  49
]);
```

# --seed--

## --seed-contents--

```js
const squareList = arr => {
  // Only change code below this line
  return arr;
  // Only change code above this line
};

const squaredIntegers = squareList([-3, 4.8, 5, 3, -3.2]);
console.log(squaredIntegers);
```

# --solutions--

```js
const squareList = arr => {
  const positiveIntegers = arr.filter(num => {
    return num >= 0 && Number.isInteger(num);
  });
  const squaredIntegers = positiveIntegers.map(num => {
    return num ** 2;
  });
  return squaredIntegers;
};
```
