---
id: 587d7da9367417b2b2512b6a
title: Rudisha safu ya data iliyopangwa bila kubadilisha safu ya data ya awali
challengeType: 1
forumTopicId: 301237
dashedName: return-a-sorted-array-without-changing-the-original-array
---

# --description--

Athari moja ya njia ya `sort` ni kwamba hubadilisha mpangilio wa vipengele katika safu ya data ya awali. Kwa maneno mengine, hubadilisha safu ya data hapo hapo. Njia moja ya kuepuka hili ni kuunganisha kwanza safu ya data tupu na ile inayopangwa (kumbuka kwamba `slice` na `concat` hurudisha safu mpya ya data), kisha tumia njia ya `sort`.

# --instructions--

Tumia njia ya `sort` katika kitendakazi cha `nonMutatingSort` kupanga vipengele vya safu ya data kwa mpangilio wa kuongezeka. Kitendakazi kinapaswa kurudisha safu mpya ya data, na kusiibadilisha thamani ya `globalArray`.

# --hints--

Msimbo wako unapaswa kutumia njia ya `sort`.

```js
assert(nonMutatingSort.toString().match(/\.sort/g));
```

Thamani ya `globalArray` haipaswi kubadilika.

```js
assert(JSON.stringify(globalArray) === JSON.stringify([5, 6, 3, 2, 9]));
```

`nonMutatingSort(globalArray)` inapaswa kurudisha `[2, 3, 5, 6, 9]`.

```js
assert(
  JSON.stringify(nonMutatingSort(globalArray)) ===
    JSON.stringify([2, 3, 5, 6, 9])
);
```

`nonMutatingSort(globalArray)` haipaswi kuwa msimbo uliowekwa moja kwa moja.

```js
assert(!nonMutatingSort.toString().match(/\[.*?[23569].*?\]/gs));
```

Kitendakazi kinapaswa kurudisha safu mpya ya data, si safu ya data iliyopitishwa.

```js
assert(nonMutatingSort(globalArray) !== globalArray);
```

`nonMutatingSort([1, 30, 4, 21, 100000])` inapaswa kurudisha `[1, 4, 21, 30, 100000]`.

```js
assert(JSON.stringify(nonMutatingSort([1, 30, 4, 21, 100000])) ===
    JSON.stringify([1, 4, 21, 30, 100000]))
```

`nonMutatingSort([140000, 104, 99])` inapaswa kurudisha `[99, 104, 140000]`.

```js
assert(JSON.stringify(nonMutatingSort([140000, 104, 99])) ===
    JSON.stringify([99, 104, 140000]))
```

# --seed--

## --seed-contents--

```js
const globalArray = [5, 6, 3, 2, 9];

function nonMutatingSort(arr) {
  // Only change code below this line


  // Only change code above this line
}

nonMutatingSort(globalArray);
```

# --solutions--

```js
const globalArray = [5, 6, 3, 2, 9];
function nonMutatingSort(arr) {
  return [].concat(arr).sort((a,b) => a-b);
}
```
