---
id: 587d7da9367417b2b2512b67
title: Ongeza vipengele mwishoni mwa safu ya data kwa kutumia concat badala ya push
challengeType: 1
forumTopicId: 301226
dashedName: add-elements-to-the-end-of-an-array-using-concat-instead-of-push
---

# --description--

Programu utendaji ni kuhusu kuunda na kutumia vitendakazi visivyo badilisha data.

Changamoto ya mwisho ilitambulisha njia ya `concat` kama njia ya kuunganisha safu za data katika safu mpya bila kubadilisha safu za data za awali. Linganisha `concat` na njia ya `push`. `push` huongeza vitu mwishoni mwa safu ile ile inayoitwa, ambayo hubadilisha safu hiyo. Hapa kuna mfano:

```js
const arr = [1, 2, 3];
arr.push(4, 5, 6);
```

`arr` ingekuwa na thamani iliyobadilishwa ya `[1, 2, 3, 4, 5, 6]`, ambayo si njia ya programu utendaji.

`concat` hutoa njia ya kuunganisha vitu vipya mwishoni mwa safu ya data bila athari za upande zinazobadilisha.

# --instructions--

Badilisha kitendakazi cha `nonMutatingPush` ili kitumie `concat` kuunganisha `newItem` mwishoni mwa `original` bila kubadilisha safu za data za `original` au `newItem`. Kitendakazi kinapaswa kurudisha safu ya data.

# --hints--

Msimbo wako unapaswa kutumia njia ya `concat`.

```js
assert(__helpers.removeJSComments(code).match(/\.concat/g));
```

Msimbo wako haupaswi kutumia njia ya `push`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?push/g));
```

Safu ya data ya `first` haipaswi kubadilika.

```js
assert(JSON.stringify(first) === JSON.stringify([1, 2, 3]));
```

Safu ya data ya `second` haipaswi kubadilika.

```js
assert(JSON.stringify(second) === JSON.stringify([4, 5]));
```

`nonMutatingPush([1, 2, 3], [4, 5])` inapaswa kurudisha `[1, 2, 3, 4, 5]`.

```js
assert(
  JSON.stringify(nonMutatingPush([1, 2, 3], [4, 5])) ===
    JSON.stringify([1, 2, 3, 4, 5])
);
```

# --seed--

## --seed-contents--

```js
function nonMutatingPush(original, newItem) {
  // Only change code below this line
  return original.push(newItem);

  // Only change code above this line
}

const first = [1, 2, 3];
const second = [4, 5];
nonMutatingPush(first, second);
```

# --solutions--

```js
function nonMutatingPush(original, newItem) {
  return original.concat(newItem);
}
const first = [1, 2, 3];
const second = [4, 5];
```
