---
id: 587d7dab367417b2b2512b6e
title: Tumia njia ya every kukagua kama kila kipengele katika safu ya data kinakidhi vigezo
challengeType: 1
forumTopicId: 301312
dashedName: use-the-every-method-to-check-that-every-element-in-an-array-meets-a-criteria
---

# --description--

Njia ya `every` hufanya kazi na safu za data ili kukagua kama *kila* kipengele kinapita mtihani fulani. Inarudisha thamani ya kweli au siyo kweli - `true` ikiwa thamani zote zinakidhi vigezo, `false` ikiwa sivyo.

Kwa mfano, msimbo ufuatao ungeangalia kama kila kipengele katika safu ya data ya `numbers` ni kidogo kuliko 10:

```js
const numbers = [1, 5, 8, 0, 10, 11];

numbers.every(function(currentValue) {
  return currentValue < 10;
});
```

Njia ya `every` ingetumia `false` hapa.

# --instructions--

Tumia njia ya `every` ndani ya kitendakazi cha `checkPositive` kukagua kama kila kipengele katika `arr` ni chanya. Kitendakazi kinapaswa kurudisha thamani ya kweli au siyo kweli.

# --hints--

Msimbo wako unapaswa kutumia njia ya `every`.

```js
assert(__helpers.removeJSComments(code).match(/\.every/g));
```

`checkPositive([1, 2, 3, -4, 5])` inapaswa kurudisha `false`.

```js
assert.isFalse(checkPositive([1, 2, 3, -4, 5]));
```

`checkPositive([1, 2, 3, 4, 5])` inapaswa kurudisha `true`.

```js
assert.isTrue(checkPositive([1, 2, 3, 4, 5]));
```

`checkPositive([1, -2, 3, -4, 5])` inapaswa kurudisha `false`.

```js
assert.isFalse(checkPositive([1, -2, 3, -4, 5]));
```

# --seed--

## --seed-contents--

```js
function checkPositive(arr) {
  // Only change code below this line


  // Only change code above this line
}

checkPositive([1, 2, 3, -4, 5]);
```

# --solutions--

```js
function checkPositive(arr) {
  return arr.every(num => num > 0);
}
```
