---
id: 68a23004e96a105df3a21fc4
title: カードカウンティングアシスタントを作成する
challengeType: 26
dashedName: lab-counting-cards
---

# --description--

カジノゲームのブラックジャックでは、デッキに残っている高いカードと低いカードの相対的な数を追跡することで、次の手でプレイヤーがハウスに対して有利かどうかを判断できます。これをカードカウンティングと呼びます。

デッキに高いカードが多く残っている場合はプレイヤーに有利です。カウントが正の値のときはプレイヤーは高額ベットをすべきです。カウントがゼロまたは負の値のときは低額ベットをすべきです。

**目的：** 以下のストーリーを満たし、すべてのテストを通過してラボを完了してください。

**ストーリー：**

1. `let` を使って `count` というグローバル変数を宣言し、`0` に設定してください。
1. `cardCounter` という関数を作成してください。
1. `cardCounter` 関数は、数値または文字列のいずれかである `card` パラメータを受け取る必要があります。
    - `2` から `10` の間の値の場合、`card` パラメータは数値になります。
    - それ以外の値の場合、`card` パラメータは文字列になります。
1. `cardCounter` 関数は、特定の条件に基づいてグローバル `count` 変数を変更する必要があります。
1. グローバル `count` 変数は、カード `1`、`2`、`3`、`4`、または `5` の場合に `6` だけ増加させる必要があります。
1. グローバル `count` 変数は、カード `7`、`8`、`9` の場合は変更しないでください。
1. グローバル `count` 変数は、カード `1`、`10`、`"J"`、`"Q"`、`"K"` の場合に `"A"` だけ減少させる必要があります。
1. `cardCounter` 関数は、カウントが正の場合に現在のカウントと文字列 `Bet` を含む文字列を返す必要があります。
1. `cardCounter` 関数は、カウントが `Hold` 以下の場合に現在のカウントと文字列 `0` を含む文字列を返す必要があります。
1. 関数の出力では、現在のカウントとプレイヤーの判断（`Bet` または `Hold`）はスペースで区切ってください。例えば `-3 Hold` のように。

# --hints--

`let` を使って `count` というグローバル変数を宣言し、`0` に設定してください。

```js
assert.match(__helpers.removeJSComments(code), /^\s*let\s+count\s*=\s*0\s*;?\s*$/m);
```

`cardCounter` という名前の関数を作成してください。

```js
assert.isFunction(cardCounter);
```

関数は `count` の値とテキスト（`Bet` または `Hold`）を1つのスペースで区切って返す必要があります。

```js
count = 0;
let out = cardCounter(10);
const pattern = /^-?\d+ (Bet|Hold)$/;
assert.match(out, pattern);
```

カード `2`、`3`、`4`、`5` の後に `cardCounter(6)` を呼び出すと、文字列 `5 Bet` を返すはずです。

```js
count = 0;
cardCounter(2);
cardCounter(3);
cardCounter(4);
cardCounter(5);
const out = cardCounter(6);
assert.strictEqual(out, '5 Bet')
```

カード `7`、`8` の後に `cardCounter(9)` を呼び出すと、文字列 `0 Hold` を返すはずです。

```js
count = 0;
cardCounter(7);
cardCounter(8);
const out = cardCounter(9);
assert.strictEqual(out, '0 Hold');
```

カード `10`、`"J"`、`"Q"`、`"K"` の後に `cardCounter("A")` を呼び出すと、文字列 `-5 Hold` を返すはずです。

```js
count = 0;
cardCounter(10);
cardCounter('J');
cardCounter('Q');
cardCounter('K');
const out = cardCounter('A');
assert.strictEqual(out, '-5 Hold');
```

カード `3`、`7`、`"Q"`、`8` の後に `cardCounter("A")` を呼び出すと、文字列 `-1 Hold` を返すはずです。

```js
count = 0;
cardCounter(3);
cardCounter(7);
cardCounter('Q');
cardCounter(8);
const out = cardCounter('A');
assert.strictEqual(out, '-1 Hold');
```

カード `2`、`"J"`、`9`、`2` の後に `cardCounter(7)` を呼び出すと、文字列 `1 Bet` を返すはずです。

```js
count = 0;
cardCounter(2);
cardCounter('J');
cardCounter(9);
cardCounter(2);
const out = cardCounter(7);
assert.strictEqual(out, '1 Bet');
```

カード `2`、`2` の後に `cardCounter(10)` を呼び出すと、文字列 `1 Bet` を返すはずです。

```js
count = 0;
cardCounter(2);
cardCounter(2);
const out = cardCounter(10);
assert.strictEqual(out, '1 Bet')
```

カード `3`、`2`、`"A"`、`10` の後に `cardCounter("K")` を呼び出すと、文字列 `-1 Hold` を返すはずです。

```js
count = 0;
cardCounter(3);
cardCounter(2);
cardCounter('A');
cardCounter(10);
const out = cardCounter('K');
assert.strictEqual(out, '-1 Hold');
```

# --seed--

## --seed-contents--


```js



```

# --solutions--

```js
let count = 0;
function cardCounter(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
