---
id: 5dc174fcf86c76b9248c6eb2
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você continuará trabalhando com elementos básicos de HTML como **títulos**, **parágrafos** e **listas** construindo um aplicativo de fotos de gatos. 

Inicie o workshop adicionando um elemento `h1` com o texto `CatPhotoApp`.

# --hints--

O texto `CatPhotoApp` deve estar presente no código.Não se esqueça de verificar a ortografia.

```js
assert.match(code, /catphotoapp/i);
```

O elemento `h1` deve ter uma tag de abertura.As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Suas tags `h1` devem estar em letras minúsculas. Por convenção, todas as tags HTML são escritas em letras minúsculas.

```js
assert.notMatch(code, /<\/?H1>/);
```

Seu elemento `h1` deve ter uma tag de fechamento. Tags de fechamento têm esta sintaxe: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

O elemento `h1` deve conter o texto `CatPhotoApp`.Você omitiu o texto, digitou o texto incorretamente ou ele não está entre as tags de abertura e fechamento do elemento `h1`.

```js
assert.equal(document.querySelector('h1')?.innerText.toLowerCase(), 'catphotoapp');
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
   
--fcc-editable-region--
  </body>
</html>
```
