---
id: 5dc1798ff86c76b9248c6eb3
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Abaixo do elemento `h1`, adicione um elemento `h2` com este texto:

`Cat Photos`

# --hints--

O elemento `h1` deve ter uma tag de abertura.As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Seu elemento `h1` deve ter uma tag de fechamento. Tags de fechamento têm esta sintaxe: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Você deve ter apenas um elemento `h1`.Remova o elemento adicional.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

O elemento `h1` deve conter o texto 'CatPhotoApp'.Você omitiu o texto ou tem um erro de digitação.

```js
assert.equal(document.querySelector('h1')?.innerText.toLowerCase(), 'catphotoapp');
```

Seus `h2` tags devem estar em letras minúsculas. Por convenção, todas as tags HTML são escritas em letras minúsculas.

```js
assert.notMatch(code, /<\/?H2>/);
```

O elemento `h2` deve ter uma tag de abertura.As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert.exists(document.querySelector('h2'));
```

O elemento `h2` deve ter uma tag de fechamento.As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert.match(code, /<\/h2\>/);
```

O texto do elemento `h2` deve ser `Cat Photos`.Coloque apenas o texto `Cat Photos` entre as tags de abertura e de fechamento de `h2`.

```js
assert.equal(document.querySelector('h2')?.innerText.toLowerCase(), 'cat photos');
```

O elemento `h2` deve estar abaixo do elemento `h1`.O elemento `h1` tem maior importância e deve estar acima do elemento `h2`.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H1'), collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>
    
--fcc-editable-region--
  </body>
</html>
```
