---
id: 66b205e6eacba4c4e54ea434
title: crea un mascheratore di email
challengeType: 26
dashedName: build-an-email-masker
---

# --description--

In questo laboratorio, maschererai la parte username di un indirizzo email con degli asterischi. Mascherare è un termine usato per nascondere o sostituire informazioni sensibili con asterischi o altri caratteri.

Per esempio, se l'indirizzo email era `myEmail@email.com`, allora l'indirizzo email mascherato sarà `m*****l@email.com`.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Crea una funzione chiamata `maskEmail` che prende `email` come argomento.
2. All'interno della funzione, dovresti mascherare `email` e aggiungere il nome del dominio. Ricorda che puoi usare metodi come `slice`, `repeat`, `indexOf` o anche `replace` per aiutarti.
3. All'esterno della funzione, dichiara una variabile chiamata `email` per archiviare l'indirizzo email che vuoi mascherare.
4. Chiama la funzione `maskEmail` con la variabile `email` e mostra il risultato nella console.
5. `maskEmail("apple.pie@example.com")` dovrebbe restituire `"a*******e@example.com"`.
6. `maskEmail("freecodecamp@example.com")` dovrebbe restituire `"f**********p@example.com"`.
7. `maskEmail("info@test.dev")` dovrebbe restituire `"i**o@test.dev"`.
8. `maskEmail("user@domain.org")` dovrebbe restituire `"u**r@domain.org"`.

# --hints--

Dovresti definire una funzione chiamata `maskEmail`.

```js
assert.isFunction(maskEmail);
```

La funzione `maskEmail` dovrebbe prendere una stringa, `email`, come argomento.

```js
assert.match(maskEmail.toString(), /\s*maskEmail\(\s*\w+\s*\)/);
```

All'esterno della funzione, dovresti avere una variabile `email`.

```js
assert.isDefined(email);
```

Dovresti assegnare un indirizzo email valido alla tua variabile `email`.

```js
assert.match(email, /^[a-zA-Z0-9._%-]+\+?[a-zA-Z0-9._%-]+@[a-zA-Z0-9.-]+\.[a-zA-Z]{2,}$/);
```

`maskEmail("apple.pie@example.com")` dovrebbe restituire `"a*******e@example.com"`.

```js
assert.strictEqual(maskEmail("apple.pie@example.com"),"a*******e@example.com");
```

`maskEmail("freecodecamp@example.com")` dovrebbe restituire `"f**********p@example.com"`.

```js
assert.strictEqual(maskEmail("freecodecamp@example.com"), "f**********p@example.com");
```

`maskEmail("info@test.dev")` dovrebbe restituire `"i**o@test.dev"`.

```js
assert.strictEqual(maskEmail("info@test.dev"), "i**o@test.dev");
```

`maskEmail("user@domain.org")` dovrebbe restituire `"u**r@domain.org"`.

```js
assert.strictEqual(maskEmail("user@domain.org"), "u**r@domain.org");
```

La tua `maskEmail` dovrebbe produrre il risultato corretto.

```js
assert.strictEqual(maskEmail("johnappleseed@email.com"), "j***********d@email.com");
assert.strictEqual(maskEmail("john.doe@example.com"), "j******e@example.com");
assert.strictEqual(maskEmail("jane.smith@sample.org"), "j********h@sample.org");
assert.strictEqual(maskEmail("michael.jordan@testmail.net"), "m************n@testmail.net");
assert.strictEqual(maskEmail("mike+nospam@testmail.net"), "m*********m@testmail.net");
```

Dovresti registrare l'output della chiamata a `maskEmail` con `email` come argomento.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(maskEmail\(email\)\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function maskEmail(email) {
  const atIndex = email.indexOf("@");
  const userName = email.slice(0, atIndex);
  const domain = email.slice(atIndex);

  const maskedName =
    userName[0] + "*".repeat(userName.length - 2) + userName[userName.length - 1] + domain;

  return maskedName;
}

const email = "apple.pie@example.com";
console.log(maskEmail(email));
```
