---
id: 587d7fb4367417b2b2512bfd
title: 给 package.json 添加关键词
challengeType: 2
forumTopicId: 301526
dashedName: add-keywords-to-your-package-json
---

# --description--

在 `keywords` 字段中可以使用相关的关键字描述项目。 这是一个示例：

```json
"keywords": [ "descriptive", "related", "words" ],
```

正如你所见的，这个字段的结构是一个由双引号字符串组成的数组。

# --instructions--

在 package.json 文件中，给 `keywords` 字段添加一个由适当的字符串组成的数组。

“freecodecamp”应该是关键词之一。

# --hints--

package.json 应该有一个有效的“keywords”键

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.keywords, '"keywords" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

“keywords”字段应该是一个数组

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.isArray(data.keywords, '"keywords" is not an array');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

“keywords”中应该包含关键词“freecodecamp”

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.include(
        data.keywords,
        'freecodecamp',
        '"keywords" does not include "freecodecamp"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
