---
id: 6995d0325fd7da40854fac40
title: 第一步
challengeType: 1
dashedName: step-1
---

# --description--

在之前的课程中，你学习了二分查找以及它与线性查找的区别。

在本次练习中，你将通过创建一个二分查找算法来实践所学内容。

首先，创建一个 `binarySearch` 函数，带有 `searchList` 和 `value` 两个参数。

`searchList` 是函数将要查找的数组，`value` 是函数将尝试在 `searchList` 中找到的目标。

# --hints--

你应该创建一个 `binarySearch` 函数。

```js
assert.isFunction(binarySearch);
```

你的 `binarySearch` 函数应该有 `searchList` 和 `value` 两个参数。

```js
const binSearchParams = __helpers.getFunctionParams(binarySearch.toString());

const binarySearchParamsNames = binSearchParams.map(param => param.name)
assert.include(binarySearchParamsNames, "searchList");
assert.include(binarySearchParamsNames, "value");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
