---
id: 6996c675b599baa7b103192d
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

在 binary search 中，有最低可能的索引和最高可能的索引。

最低可能的索引表现为当前被查找值部分中的最左边界，在本例中是 `searchList`。另一方面，最高可能的索引表现为被查找值中的最右边界。

为了考虑这两个边界，你必须考虑正在查找的整个数组。所以，在 `binarySearch` 函数内部，声明一个值为 `0` 的 `low` 变量，以及一个值为数组最后一个索引的 `high` 变量。

# --hints--

你应该在 `binarySearch` 函数中声明一个名为 `low` 的变量，并将其设置为 `0`。

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { low } = binarySearch?.variables ?? {};
assert.isTrue(low?.value.matches("0"));
```

你应该在 `binarySearch` 函数中声明一个名为 `high` 的变量，并将其设置为 `searchList.length - 1`。

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { high } = binarySearch?.variables ?? {};
assert.isTrue(high?.value.matches("searchList.length - 1"));
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
