---
id: 6996cb5869acf58de751ae81
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

回想一下，该函数接收一个名为 `searchList` 的数组和你要查找的 `value`。

现在，使用 `binarySearch([1, 2, 3, 4, 5], 3)` 和 `binarySearch([1, 2, 3, 4, 5, 9], 4)` 调用该函数，并将两个语句都记录到控制台。

# --hints--

你应该使用 `binarySearch([1, 2, 3, 4, 5], 3)` 调用该函数，并将结果记录到控制台。

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*2\s*,\s*3\s*,\s*4\s*,\s*5\s*\]\s*,\s*3\s*\)\s*\)/
);
```

你应该使用 `binarySearch([1, 2, 3, 4, 5, 9], 4)` 调用该函数，并将结果记录到控制台。

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*2\s*,\s*3\s*,\s*4\s*,\s*5\s*,\s*9\s*\]\s*,\s*4\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
    break;
  }
  return [];
}

--fcc-editable-region--  

--fcc-editable-region--  
```
