---
id: 6996c60281feca6718cc15c1
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

在你的 `binarySearch` 函数内部，创建一个 `pathToTarget` 变量并将其初始化为空数组，作为开始。

# --hints--

你应该在 `pathToTarget` 函数内有一个 `[]` 变量被设置为空数组 `binarySearch`。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+pathToTarget\s*=\s*\[\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
