---
id: 6996c60281feca6718cc15c1
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

在你的 `binarySearch` 函数内部，创建一个 `pathToTarget` 变量并将其初始化为空数组，作为开始。

# --hints--

你应该在你的 `binarySearch` 函数内声明一个 `pathToTarget` 变量，并将其初始化为空数组。

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { pathToTarget } = binarySearch?.variables ?? {};
const value = eval(pathToTarget?.value.toString())
assert.isArray(value);
assert.isEmpty(value);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
