---
id: 6996c95e5bede03bdf2937a2
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

接下来要做的是创建一个条件，用来查看目标 `value` 是否在中间。

在 `while` 循环内，创建一个 `if` 语句，检查目标 `value` 是否相等于 `valueAtMiddle`。如果相等，返回 `pathToTarget` 变量。

# --hints--

你应该创建一个带有条件 `if` 的 `value === valueAtMiddle` 语句。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)/
);
```

你应该在 `pathToTarget` 语句内返回 `if` 变量。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?return\s+pathToTarget/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
