---
id: 6996ccadf9048c79b857bb15
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

你可以看到第一次调用返回了 `[3]`。这是因为 `3` 是 `[1, 2, 3, 4, 5]` 中间的值。

另一方面，第二次调用返回了一个空数组。这是因为，在检查中间值之后，循环目前缺少通过更新 `low` 或 `high` 变量来缩小查找范围的逻辑。它只是作为一次中点检查，然后结束，没有进一步查找。

为了让二分查找继续缩小查找范围，添加一个 `else if` 语句，检查 `value` 是否大于 `valueAtMiddle`。

# --hints--

你应该添加一个 `else if` 语句，条件是 `value > valueAtMiddle`。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s+if\s*\(\s*value\s*>\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    if (value === valueAtMiddle) {
      return pathToTarget;
    } 
--fcc-editable-region--
    break;
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
