---
id: 6996c8f94fe6b250f11a2a1e
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

现在，将 `valueAtMiddle` 添加到 `pathToTarget` 数组中，以跟踪搜索过程中所经过的步骤，无论它是否是目标。你可以使用 `push()` 方法来完成这项操作。

# --hints--

你应该使用 `push()` 方法将 `valueAtMiddle` 添加到 `pathToTarget` 数组中。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /pathToTarget\s*\.\s*push\(\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
