---
id: 6996c8375f4301004459ee85
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

现在你已经计算出了 `mid` 索引，需要从 `searchList` 中实际取出该索引处的值。这个值就是你将用来与 `value` 参数（你正在查找的目标）进行比较的内容。

创建一个 `valueAtMiddle` 变量，并将其赋值为位于 `searchList` 索引处的 `mid` 元素。

# --hints--

你应该创建一个 `valueAtMiddle` 变量，并将其设置为 `searchList[mid]`。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+valueAtMiddle\s*=\s*searchList\s*\[\s*mid\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
