---
id: 6996c675b599baa7b103192d
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

在二分查找中，有最低可能的索引和最高可能的索引。

最低可能的索引表现为当前被查找值部分中的最左边界，在本例中是 `searchList`。另一方面，最高可能的索引表现为被查找值中的最右边界。

为了考虑这两个边界，你必须考虑整个被查找的数组。因此，在 `binarySearch` 函数内部，定义一个 `low` 变量，其值为 `0`，以及一个 `high` 变量，其值考虑了被查找数组中的最后一个索引。

# --hints--

你应该在 `low` 函数中有一个名为 `0` 的变量，且其值为 `binarySearch`。

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /(const|let|var)\s+low\s*=\s*0/);
```

你应该在 `high` 函数中有一个名为 `searchList.length - 1` 的变量，且其值为 `binarySearch`。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+high\s*=\s*searchList\.length\s*-\s*1/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
