---
id: 6996d1032feed88e2957dc8e
title: 第 15 步
challengeType: 1
dashedName: step-15
---

# --description--

你可以看到最后一次调用打印了一个空数组。为了更明确地表示该值不在被查找的数组中，更新 `return` 语句，使其返回一个包含两个元素的数组：一个空数组和字串 `Value not found`。

# --hints--

你应该更新 `return` 语句，使函数返回一个数组，第一元素是 `[]`，第二元素是 `"Value not found"`。

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const validExpressions = [
  "[[], 'Value not found']",
  "[[], \"Value not found\"]",
  "[[], `Value not found`]"
];
assert.isTrue(validExpressions.some(expr => binarySearch?.hasReturn(expr) || binarySearch?.hasReturn("(" + expr + ")")));
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
--fcc-editable-region--
  return [];
--fcc-editable-region--
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```
