---
id: 6996ca21a74104721f2c7e5d
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

你需要测试内容，这样你才能在这个初始阶段理解算法的流程。

为此，你需要先跳出 loop。这是因为当前的 implementation 只允许一次 iteration，所以如果条件不满足，就会出现 无限循环。

就在 `if` 语句之后，使用 `break` 关键字跳出 `while` 循环。然后，在 `while` 循环之后，返回一个空数组以表示未找到该值。

# --hints--

你应该在 `if` 语句后使用 `break` 关键字跳出 `while` 循环。

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearch?.toString());
assert.match(
  codeWithoutComments,
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?\}\s*break\s*;?/
);
```

你应该在 `while` 循环之后返回一个空数组。

```js
assert.deepEqual(binarySearch([], NaN), []);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--
    
  }
  
--fcc-editable-region--
}
```
