---
id: 6996c7363a4a437839f1f15d
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

既然你已经定义了查找的边界，现在是时候创建执行二分查找算法的循环了。

二分查找通过不断缩小查找范围来工作。只要存在有效的元素范围需要查看，这个过程就会持续。

你可以用一个 `while` 循环来表达，只要你的 `low` 指针小于或等于 `high` 指针，这个循环就会继续。

如果 `low` 变得大于 `high`，这意味着查找范围已经为空，`value` 不在数组中。

在 `binarySearch` 函数内部，创建一个 `while` 循环，条件是检查 `low` 是否小于或等于 high。

# --hints--

你应该创建一个 `while` 循环，当 `low` 小于或等于 `high` 时运行。

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /while\s*\(\s*low\s*<=\s*high\s*\)/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
