---
id: 6996ca21a74104721f2c7e5d
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

你需要测试一些内容，这样你才能理解算法在这个初始阶段的流程。

为此，你首先需要跳出循环。因为当前的实现只允许一次迭代，如果条件不满足，就会出现无限循环。

就在 `if` 语句之后，使用 `break` 关键字跳出 `while` 循环。然后，在 `while` 循环之后，返回一个空数组以表示未找到该值。

# --hints--

你应该在 `break` 语句之后使用 `while` 关键字跳出 `if` 循环。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?\}\s*break\s*;?/
);
```

你应该在 `while` 循环之后返回一个空数组。

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /return\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--
    
  }
  
--fcc-editable-region--
}
```
