---
id: 6996c796385b6ea4e222d5e8
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

要开始循环，你需要查找当前搜索空间的中间索引。

你可以通过计算 `low` 和 `high` 变量的平均值来获得中间索引。你必须使用向下取整除法来计算两个值相加后的平均值，这样答案总是向下取整的整数，因为索引必须是整数。

在你的 `while` 循环内，创建一个 `mid` 变量，设置为 `low` 和 `high` 的平均值。

# --hints--

在 while 循环内，你应该创建一个 `mid` 变量，并将其设置为 `Math.floor((low + high) / 2)`。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+mid\s*=\s*Math\.floor\(\s*\(\s*low\s*\+\s*high\s*\)\s*\/\s*2\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
--fcc-editable-region--
      
--fcc-editable-region--
  }
}
```
