---
id: 6996cfb6c01e9a802add38be
title: 第 13 步
challengeType: 1
dashedName: step-13
---

# --description--

最后，如果 `value` 没有在中间找到，并且也不大于 `valueAtMiddle`，那么 `value` 必须小于 `valueAtMiddle`，并且必须在左边。最终的 `else` 块将处理这个情况。

所以，通过移除 `break` 关键字并添加一个 `else` 块来完成循环。在其中，通过从 `high` 中减去 `1` 来更新 `mid` 变量。

# --hints--

你应该用一个 `break` 块替换 `else` 关键字，该块将 `high` 变量更新为 `mid - 1`。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s*\{\s*high\s*=\s*mid\s*-\s*1\s*;?\s*\}/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    }

--fcc-editable-region--      
    break;
--fcc-editable-region--

  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
