---
id: 6996d0a2c2ac55041145a13f
title: 第 14 步
challengeType: 1
dashedName: step-14
---

# --description--

你注意到第二次函数调用时列表中有 `3, 5, 4` 这些值了吗？这表明二分查找算法正在按预期工作。

事情是这样发生的：算法首先将 `3` 作为初始列表的中间值进行检查。由于 `4` 大于 `3`，查找范围向右半部分移动。接着它将 `5` 作为新的中间值进行检查。因为 `4` 小于 `5`，查找范围向左移动，最终确定 `4` 是最终范围的中间值。

要再次测试这个函数，使用 `[1, 3, 5, 9, 14, 22], 10` 调用它并将结果记录到控制台。这是一个找不到该值的情况。

# --hints--

你应该用 `binarySearch([1, 3, 5, 9, 14, 22], 10)` 调用函数并将结果记录到控制台。

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*3\s*,\s*5\s*,\s*9\s*,\s*14\s*,\s*22\s*\]\s*,\s*10\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
--fcc-editable-region--

--fcc-editable-region--
```
