---
id: 69a5f35669099ed52f8563b1
title: スマートパントリー補充プログラムを作成する
challengeType: 26
dashedName: lab-smart-pantry-restocker
---

# --description--

このラボでは、配列、オブジェクト、ループ、条件分岐などの基本的なJavaScriptの概念を使って、小さなパントリー管理プログラムを作成します。

パントリーに届いた商品の入荷をシミュレートし、各アイテムに対して何をするか決め、保管のために結果を整理します。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストを通過してラボを完了してください。

**ユーザーストーリー:**

1. 文字列の配列を受け取り、`parseShipment(rawData)` プロパティを持つオブジェクトの配列を返す`{ sku, name, qty, expires, zone }` 関数を実装してください。

2. 現在のパントリーと入荷品を比較し、`planRestock(pantry, shipment)` 形式のアクションの配列を返す`{ type: "restock" | "discard" | "donate", item }` 関数を実装してください。

3. 各アイテムの`groupByZone(actions)` プロパティに基づいてアクションを保管ゾーンごとにグループ化する`zone` 関数を実装してください。

4. パントリーのディープコピーを返す`clonePantry(pantry)` 関数を実装してください。ディープコピーとは、新しい配列と新しいオブジェクトを作成し、コピーを変更しても元のパントリーが変わらないようにすることです。

5. すべての関数を組み合わせて入荷品を処理し、最終的にグループ化された結果オブジェクトをコンソールに出力してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`parseShipment` という文字列の配列をパラメータに受け取る`rawData` という名前の関数を定義してください。

``` js
assert.isFunction(parseShipment);
assert.lengthOf(parseShipment, 1);
```

`parseShipment` 関数は、配列内の入荷品の文字列を、`sku`、`name`、`qty`、`expires`、`zone` のプロパティを持つオブジェクトに変換する必要があります。

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);

assert.isArray(result);
assert.isObject(result[0]);
assert.deepEqual(result[0], {
  sku: "A10",
  name: "Tomatoes",
  qty: 5,
  expires: "2027-01-01",
  zone: "fridge"
});
```

入荷品のSKUが重複している場合は無視してください。

``` js
const rawData = [
  "C32|Eggs|3|2027-01-01",
  "C32|Eggs|3|2027-01-01"
];

const result = parseShipment(rawData);

assert.lengthOf(result, 1);
assert.strictEqual(result[0].sku, "C32");
```

入荷品の文字列に`zone` が含まれていない場合は、`general` を代入してください。

``` js
const rawData = ["B21|Bananas|10|2027-01-01"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].zone, "general");
```

`qty` の値は数値に変換してください。

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].qty, 5);
assert.isNumber(result[0].qty);
```

`planRestock` と `pantry` の2つのパラメータを受け取る`shipment` という名前の関数を定義してください。

``` js
assert.isFunction(planRestock);
assert.lengthOf(planRestock, 2);
```

`planRestock` 関数は、`type` と `item` のプロパティを持つアクションの配列を返す必要があります。

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 6, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.isArray(result);
assert.isObject(result[0]);
assert.property(result[0], "type");
assert.property(result[0], "item");
```

入荷品の数量が`0` 以下の場合は、パントリーにあるかどうかに関わらず、タイプが`discard` のアクションを作成してください。

``` js
const pantry = [{ sku: "D43", name: "Pineapples", qty: 2, expires: "2027-01-01", zone: "general" }];
const shipment = [{ sku: "D43", name: "Pineapples", qty: 0, expires: "2027-01-01", zone: "general" },
                  { sku: "E54", name: "Peppers", qty: -1, expires: "2027-01-01", zone: "fridge" }
                ];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "discard");
assert.strictEqual(result[1].type, "discard");
```

入荷品がすでにパントリーに存在する場合（同じ`sku`）、アクションは`restock` にしてください。

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "restock");
```

入荷品がパントリーに存在しない場合は、アクションは`"donate"` にしてください。

``` js
const pantry = [];
const shipment = [{ sku: "C32", name: "Eggs", qty: 3, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "donate");
```

`groupByZone` というパラメータを受け取る`actions` という名前の関数を定義してください。

``` js
assert.isFunction(groupByZone);
assert.lengthOf(groupByZone, 1);
```

`groupByZone` 関数は、各オブジェクトの`zone` プロパティでアクションをグループ化して返す必要があります。

``` js

const actions = [
  { type: "restock", item: { sku: "A1", zone: "fridge" } },
  { type: "restock", item: { sku: "B1", zone: "pantry" } }, 
];

const result = groupByZone(actions);

assert.isObject(result);
assert.isArray(result.fridge);
assert.isArray(result.pantry);
```

`clonePantry` というパラメータを受け取る`pantry` という名前の関数を定義してください。

``` js
assert.isFunction(clonePantry);
assert.lengthOf(clonePantry, 1);
```

`clonePantry` は元のパントリー配列ではなく、新しい配列を返す必要があります。

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.isArray(pantryClone);
assert.notStrictEqual(pantryClone, pantry);
```

複製されたパントリー内のオブジェクトも新しいオブジェクト（ディープコピー）である必要があります。

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.notStrictEqual(pantryClone[0], pantry[0]);
assert.deepEqual(pantryClone[0], pantry[0]);

pantryClone[0].qty = 100;
assert.strictEqual(pantry[0].qty, 4);
```

関数は連携して入荷品を処理し、結果のアクションをグループ化する必要があります。

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];

const shipment = parseShipment(rawData);
const pantryClone = clonePantry(pantry);
const actions = planRestock(pantryClone, shipment);
const grouped = groupByZone(actions);

assert.isObject(grouped);
assert.strictEqual(grouped.fridge[0].item.qty, 5);
```

ゾーンごとにグループ化された結果のアクションをコンソールに出力してください。

``` js
const groupedLog = getLogs().find(log => {
  if (!log || typeof log !== "object" || Array.isArray(log)) return false;

  return Object.values(log).some(zoneActions => {
    if (!Array.isArray(zoneActions)) return false;

    return zoneActions.some(action =>
      action &&
      ["restock", "discard", "donate"].includes(action.type) &&
      action.item &&
      typeof action.item.sku === "string"
    );
  });
});

assert.isOk(groupedLog, "Log the grouped actions object to the console.");
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js

function parseShipment(rawData) {

  // Early return if rawData is not an Array
  if (!Array.isArray(rawData)) return [];

  const result = [];
  const seenSKUs = [];

  for (let line of rawData) {

    const parts = line.split("|");

    const sku = parts[0].trim();

    if (seenSKUs.includes(sku)) continue; // Avoid duplicates

    const name = parts[1].trim();
    const qty = Number(parts[2].trim()); // Parse qty into a number
    const expires = parts[3].trim();
    const zone = ( parts[4] ) ? parts[4].trim() : "general";

    result.push({
      sku: sku,
      name: name,
      qty: qty,
      expires: expires,
      zone: zone
    });

    seenSKUs.push(sku);

  }

  return result;
}

function planRestock(pantry, shipment) {
  const actions = [];

  for (let item of shipment) {

    var exists = false;

    for (let stock of pantry)
    {
      if (stock.sku === item.sku)
      {
        exists = true;
      }
    }

    if (item.qty > 0) {
      if (exists) {
        actions.push({ type: "restock", item: item });
      } else{
        actions.push({ type: "donate", item: item });
      }
    } else {
      actions.push({ type: "discard", item: item });
    }
  }

  return actions;
}

function groupByZone(actions) {

  const byZone = {};

  for (const action of actions) {

    const zone = action.item.zone;

    if (!byZone[zone]) {
      byZone[zone] = [];
    }

    byZone[zone].push(action);
  }

  return byZone;

}

function clonePantry(pantry) {
  const pantryClone = [];

  for (let item of pantry) {
    pantryClone.push({
      sku: item.sku,
      name: item.name,
      qty: item.qty,
      expires: item.expires,
      zone: item.zone
    });
  }

  return pantryClone;
}

const pantry = [
  { sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" },
  { sku: "D43", name: "Pineapples", qty: 2, expires: "2020-01-01", zone: "general" }
];

const rawData = [
  "A10|Tomatoes|5|2027-01-01", // Restock existing item
  "B21|Bananas|10|2027-01-01", // Donate new item without zone
  "C32|Eggs|3|2027-01-01|fridge", // Donate to a defined zone
  "C32|Eggs|3|2027-01-01", // Duplicated SKU in shipment
  "D43|Pineapples|0|2027-01-01", // Discard with quantity of 0
  "E54|Peppers|-1|2027-01-01|fridge" // Discard even if it's not in pantry
];

const shipment = parseShipment(rawData);
const pantryCopy = clonePantry(pantry);
const actions = planRestock(pantryCopy, shipment);
const grouped = groupByZone(actions);

// Show items expanded in zones
console.log(grouped);

```
