---
id: 587d8259367417b2b2512c86
title: تنفيذ فرز الإدراج
challengeType: 26
dashedName: implement-insertion-sort
---

# --description--

في هذا المختبر، ستنفذ خوارزمية فرز الإدراج. تعمل هذه الطريقة عن طريق بناء مصفوفة مرتبة في بداية القائمة. تبدأ المصفوفة المرتبة بالعنصر الأول. ثم تفحص العنصر التالي وتبدله إلى الخلف داخل المصفوفة المرتبة حتى يصبح في الموضع المرتب. تستمر في التكرار عبر القائمة وتبديل العناصر الجديدة إلى الخلف داخل الجزء المرتب حتى تصل إلى النهاية. هذه الخوارزمية لها تعقيد زمني تربيعي في الحالات المتوسطة والأسوأ.

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ دالة `insertionSort`.
1. يجب أن تأخذ دالة `insertionSort` مصفوفة من الأعداد الصحيحة وتُرجع مصفوفة تحتوي على نفس الأعداد مرتبة من الأصغر إلى الأكبر.

# --hints--

يجب أن تكون `insertionSort` دالة.

```js
assert.isFunction(insertionSort);
```

يجب أن تُرجع `insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` مصفوفة لم تتغير إلا في الترتيب.

```js
assert.sameMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

يجب أن تُرجع `insertionSort` مصفوفة مرتبة (من الأصغر إلى الأكبر).

```js

assert.sameOrderedMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

يجب أن تُرجع `insertionSort([5, 4, 33, 2, 8])` `[2, 4, 5, 8, 33]`.

```js
assert.sameOrderedMembers(insertionSort([5, 4, 33, 2, 8]), [2, 4, 5, 8, 33])
```

يجب ألا تستخدم `insertionSort` الطريقة المدمجة `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    insertionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function insertionSort (array) {
  for (let currentIndex = 0; currentIndex < array.length; currentIndex++) {
    let current = array[currentIndex];
    let j = currentIndex - 1;
    while (j > -1 && array[j] > current) {
      array[j + 1] = array[j];
      j--;
    }
    array[j + 1] = current;
  }
  return array;
}
```
