---
id: 68e4c81577243e4db1676491
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você praticará trabalhar com HTML semântico construindo uma página web que inclui algumas das dicas de Quincy Larson para conseguir um emprego como desenvolvedor. O boilerplate básico de HTML foi preparado para você.

Comece criando um elemento `h1` com o texto `Quincy's Tips for Getting a Developer Job`.

# --hints--

Você deve ter um elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Seu elemento `h1` deve conter o texto `Quincy's Tips for Getting a Developer Job`.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), "Quincy's Tips for Getting a Developer Job")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
--fcc-editable-region--
   
--fcc-editable-region--
  </body>
</html>
```
