---
id: 68e65ec5348e40a739903e6d
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Agora, crie um elemento de parágrafo abaixo do elemento `h1`. Dentro deste parágrafo adicione o texto `Learning to code is hard, but as Quincy Larson says, You can become a developer.`.

# --hints--

Você deve ter um elemento `p`.

```js
assert.exists(document.querySelector('p'));
```

Seu parágrafo deve conter o texto `Learning to code is hard, but as Quincy Larson says, You can become a developer.`. Verifique sua ortografia cuidadosamente.

```js
assert.equal(document.querySelector('p')?.textContent.trim(), "Learning to code is hard, but as Quincy Larson says, You can become a developer.");
```

Seu parágrafo deve estar abaixo do elemento `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
