---
id: 68e964b84f954c322376e592
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

If the source of a quote is a website, you can cite it with the `cite` attribute. O valor deste atributo deve ser uma URL válida. Embora este atributo não altere a apresentação do bloco de citação, ele é muito útil para fornecer aos leitores de tela e motores de busca mais informações sobre a citação.

Aqui está um exemplo de um elemento de citação inline com um atributo `cite`:

```html
<p>
  Nancy said,
  <q cite="https://example.com">Learning is fun!</q>
</p>
```

Adicione o atributo `cite` ao elemento de citação inline com esta URL:

`https://www.freecodecamp.org/news/learn-to-code-book/`

# --hints--

Seu elemento `q` deve ter um atributo `cite`.

```js
const qEl = document.querySelector('h1 + p > q');
assert.exists(qEl?.getAttribute('cite'));
```

O atributo `cite` do seu elemento `q` deve ter o valor `https://www.freecodecamp.org/news/learn-to-code-book/`.

```js
const qEl = document.querySelector('h1 + p > q');
assert.equal(qEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
--fcc-editable-region--
      <q>You can become a developer.</q>
--fcc-editable-region--
    </p>
  </body>
</html>
```
