---
id: 68e90335bd53a901c41dfc13
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Em uma lição anterior você aprendeu que o elemento `q` é usado para distinguir texto citado do conteúdo ao redor.

Aqui está um exemplo:

```html
<p>
  Nancy said, <q>Learning is fun!</q>
</p>
```

A maioria dos navegadores modernos adicionará aspas automaticamente em torno de uma citação inline quando você usar o elemento `q`.

Dentro do parágrafo, envolva `You can become a developer.` em um elemento de citação inline, mantendo o resto do parágrafo inalterado.

# --hints--

Você deve ter um elemento `q`.

```js
assert.exists(document.querySelector('q'));
```

Seu elemento `q` deve estar aninhado dentro do parágrafo que está abaixo do elemento `h1`.

```js
assert.exists(document.querySelector('h1 + p > q'));
```

Seu elemento `q` deve conter o texto `You can become a developer.`

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent.trim(), 'You can become a developer.');
```

Dentro do seu elemento `q`, o texto não deve ter espaços em branco no início ou no fim (sem espaços, tabulações ou quebras de linha antes do primeiro caractere ou depois do último).

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent, 'You can become a developer.');
```

Você deve adicionar `q` tags apenas na parte apropriada da frase sem alterar o texto original. O elemento `p` ainda deve conter: `Learning to code is hard, but as Quincy Larson says, You can become a developer.`

```js
const pElCleanedText = document.querySelector('h1 + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, 'Learning to code is hard, but as Quincy Larson says, You can become a developer.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    <p>Learning to code is hard, but as Quincy Larson says, You can become a developer.</p>
--fcc-editable-region--
  </body>
</html>
```
