---
id: 68ea7e7b916e8270483689b6
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Abaixo do elemento `h2`, você adicionará outra citação de Quincy. Desta vez, a citação não fará parte de um parágrafo maior. Em vez disso, o parágrafo inteiro será uma citação. Para distinguir um texto citado assim, você deve usar o elemento de citação em bloco: `blockquote`. No navegador, você verá que o texto fica ligeiramente indentado.

Aqui está um exemplo de um elemento de citação em bloco com texto citado:

```html
<blockquote>
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

Agora, dentro da primeira seção, adicione um elemento de citação em bloco abaixo do elemento `h2` com o texto `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`.

# --hints--

Dentro do primeiro elemento `section`, você deve ter um elemento `blockquote` abaixo do elemento `h2`.

```js
assert.exists(document.querySelector('section:first-of-type > h2 + blockquote'));
```

Seu elemento `blockquote` deve conter o texto `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.innerText.trim(), 'Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/"
        >You can become a developer.<
      /q>
    </p>

    <main>
--fcc-editable-region--
      <section>
        <h2>Envisioning Success</h2>

      </section>
--fcc-editable-region--
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
