---
id: 68eaa9ef8a0867071f195f59
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Até agora você tem usado o atributo `cite` para atribuir a fonte da citação, mas o atributo não mostra realmente a fonte para o usuário.

Se você quiser atribuir a fonte visualmente, pode adicionar um elemento de citação, `cite`, fora do elemento de citação em bloco. O elemento de citação é um elemento HTML que você pode usar para marcar o título de uma obra criativa referenciada, como um livro, artigo, música, filme, site ou artigo de pesquisa.

Aqui está um exemplo:

```html
<div>
  <blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
    The first thing you should consider about education is this is an economic decision.
  </blockquote>
  <p>—Quincy Larson, <cite>Is College Still Worth it? Tips from my 20 Years in Adult Education</cite></p>
</div>
```

Dentro do elemento `p` abaixo do elemento de citação em bloco, envolva `How to Learn to Code and Get a Developer Job [Full Book]` em um elemento `cite`.

# --hints--

Você deve ter um elemento `cite` dentro do elemento `p` abaixo do elemento `blockquote`.

```js
assert.exists(document.querySelector('blockquote + p > cite'));
```

Seu elemento `cite` deve conter o texto `How to Learn to Code and Get a Developer Job [Full Book]`.

```js
const citeEl = document.querySelector('blockquote + p > cite');
assert.equal(citeEl?.innerText.trim(), 'How to Learn to Code and Get a Developer Job [Full Book]');
```

Ignorando o elemento `cite` que você acabou de adicionar, o texto do elemento `p` ainda deve ser `—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]`.

```js
const pElCleanedText = document.querySelector('blockquote + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, '—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
--fcc-editable-region--
        <p>
          —Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]
        </p>
--fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
