---
id: 68ea9d23318b6481ac3de777
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Exatamente como no elemento de citação inline, você também pode adicionar um atributo `cite` a um elemento de citação em bloco.

Aqui está um exemplo de um elemento de citação em bloco com um atributo `cite`:

```html
<blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

Agora, adicione um atributo `cite` ao elemento de citação em bloco com a URL `https://www.freecodecamp.org/news/learn-to-code-book/`.

# --hints--

Seu elemento `blockquote` deve ter um atributo `cite`.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.exists(firstBlockquoteEl?.getAttribute('cite'));
```

O atributo `cite` deve ter o valor `https://www.freecodecamp.org/news/learn-to-code-book/`.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/"
        >You can become a developer.<
      /q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
--fcc-editable-region--
        <blockquote>
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
--fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
