---
id: 68c341332d98fe38fae733af
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Ahora es momento de empezar a construir la tabla. 

Comienza agregando un elemento `table` debajo de tu elemento `h1`. Dentro del elemento `table`, agrega un elemento `caption` con el texto `Schedule by Track and Time`.

# --hints--

Debes tener un elemento `table`.

```js
assert.exists(document.querySelector("table"));
```

Tu elemento `table` debe tener un elemento `caption`.

```js
assert.exists(document.querySelector("table caption"));
```

Tu elemento `caption` debe tener el texto `Schedule by Track and Time`.

```js
assert.equal(document.querySelector("table caption")?.textContent.trim(), "Schedule by Track and Time");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

--fcc-editable-region--
    
--fcc-editable-region--
</body>
</html>
```
