---
id: 68c34ace730ca43e5fa7b8d6
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Ahora es momento de agregar otra fila a la tabla. 

Comienza agregando otro elemento `tr`. Dentro de ese elemento `tr`, agrega un elemento `th` con un atributo `scope` establecido en `"row"` y el contenido de texto `10:00 AM`. 

Luego, agrega tres elementos `td` con el siguiente contenido de texto:

- `Accessibility Deep Dive`
- `CSS for Beginners`
- `Inclusive Design Principles`

# --hints--

Debes tener un segundo elemento `tr` dentro del elemento `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 2);
```

Debes tener un elemento `th` dentro del segundo elemento `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[1].querySelector("th"));
```

Tu elemento `th` debe tener el contenido de texto `10:00 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.textContent.trim(), "10:00 AM");
```

Tu elemento `th` debe tener un atributo `scope` establecido en `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.getAttribute("scope"), "row");
```

Tu segundo elemento `tr` debe tener tres elementos `td`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[1].querySelectorAll("td"), 3);
```

Tu primer elemento `td` debe tener el contenido de texto `Accessibility Deep Dive`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[0]?.textContent.trim(), "Accessibility Deep Dive");
```

Tu segundo elemento `td` debe tener el contenido de texto `CSS for Beginners`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[1]?.textContent.trim(), "CSS for Beginners");
```   

Tu tercer elemento `td` debe tener el contenido de texto `Inclusive Design Principles`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[2]?.textContent.trim(), "Inclusive Design Principles");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
