---
id: 68c33e502cd97f2fd52cb63b
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, aprenderás cómo hacer que los elementos de la tabla sean accesibles construyendo un horario para una conferencia tecnológica.

Comienza agregando un elemento `h1` con el texto `Tech Conference 2025 Schedule`.

# --hints--

Debería tener un elemento `h1`.

```js
const h1 = document.querySelector("h1");
assert.exists(h1);
```

Tu elemento `h1` debe tener el texto `Tech Conference 2025 Schedule`.

```js
const h1 = document.querySelector("h1");
assert.strictEqual(h1?.textContent.trim(), "Tech Conference 2025 Schedule");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
    --fcc-editable-region--
    
    --fcc-editable-region--
</body>
</html>
```
