---
id: 68c3482278a14e3cf225e897
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Dentro de tu elemento `tr`, agrega un elemento `th` con el texto `9:00 AM`. Luego, debajo de ese elemento `th`, agrega tres elementos `td` con los textos `Keynote: Tech Future`, `Intro to Web Dev` y `UX for All`. 

# --hints--

Debes tener un elemento `th` dentro de tu elemento `tr`.

```js
assert.exists(document.querySelector("tbody tr th"));
```

Tu elemento `th` debe tener el contenido de texto `9:00 AM`.

```js
assert.equal(document.querySelector("tbody tr th")?.textContent.trim(), "9:00 AM");
```

Debes tener un total de tres elementos `td` debajo de tu elemento `th`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.lengthOf(tdElements, 3);
```

Tu primer elemento `td` debe tener el contenido de texto `Keynote: Tech Future`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[0]?.textContent.trim(), "Keynote: Tech Future");
```

Tu segundo elemento `td` debe tener el contenido de texto `Intro to Web Dev`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[1]?.textContent.trim(), "Intro to Web Dev");
``` 

Tu tercer elemento `td` debe tener el contenido de texto `UX for All`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[2]?.textContent.trim(), "UX for All");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
