---
id: 68c346d03709343bff132707
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Para los siguientes pasos, construirás el cuerpo de la tabla.

Comienza agregando un elemento `tbody` debajo del elemento `thead`. Luego, dentro del `tbody`, agrega un elemento `tr`.

# --hints--

Debes tener un elemento `tbody` debajo del elemento `thead`.

```js
assert.exists(document.querySelector("tbody"));
```

Debes tener un elemento `tr` dentro del elemento `tbody`.

```js
assert.exists(document.querySelector("tbody tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>
    
    --fcc-editable-region--
    
    --fcc-editable-region--
  </table>
</body>
</html>
```
