---
id: 68c3420cb0176439a95af1b8
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

El siguiente paso es agregar el encabezado de la tabla. 

Comienza agregando un elemento `thead` debajo del elemento `caption`. Dentro del `thead`, agrega un elemento `tr`. 

# --hints--

Debes tener un elemento `thead`.

```js
assert.exists(document.querySelector("table thead"));
```

Debes tener un elemento `tr` dentro del `thead`.

```js
assert.exists(document.querySelector("table thead tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    --fcc-editable-region--
        
    --fcc-editable-region--
  </table>
</body>
</html>
```
