---
id: 68c34329d430583a6cd5f7f1
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Dentro de tu elemento `tr`, agrega cuatro elementos `th`. El primer elemento `th` debe tener el texto `Time`, el segundo debe tener el texto `Track A`, el tercero debe tener el texto `Track B` y el cuarto debe tener el texto `Track C`.

# --hints--

Tu elemento `tr` debe tener cuatro elementos `th`.

```js
assert.lengthOf(document.querySelectorAll("tr > th"), 4);
```

Tu primer elemento `th` debe tener el texto `Time`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(1)")?.textContent.trim(), "Time");
```

Tu segundo elemento `th` debe tener el texto de `Track A`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(2)")?.textContent.trim(), "Track A");
```

Tu tercer elemento `th` debe tener el texto de `Track B`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(3)")?.textContent.trim(), "Track B");
```

Tu cuarto elemento `th` debe tener el texto de `Track C`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(4)")?.textContent.trim(), "Track C");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

--fcc-editable-region--
    <thead>
      <tr>
        
      </tr>
    </thead>
--fcc-editable-region--
  </table>
</body>
</html>
```
