---
id: 68c34f4aeb2b1a4170506751
title: Paso 13
challengeType: 0
dashedName: step-13
---

# --description--

A continuación, agrega una quinta fila a la tabla. Comienza añadiendo otro elemento `tr`. Dentro de ese elemento `tr`, agrega un elemento `th` con un atributo `scope` establecido en `"row"` y el contenido de texto `12:30 PM`. Luego, debajo de ese elemento `th`, agrega un elemento `td` con el contenido de texto `Lunch Break`. Tu elemento `td` también debe tener un atributo `colspan` establecido en `3` para que abarque las tres pistas.

# --hints--

Debes tener un quinto elemento `tr` dentro del elemento `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 5);
```

Debes tener un elemento `th` dentro del quinto elemento `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("th"));
```

Tu elemento `th` debe tener el contenido de texto `12:30 PM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.textContent.trim(), "12:30 PM");
```

Tu elemento `th` debe tener un atributo `scope` establecido en `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.getAttribute("scope"), "row");
```

Debes tener un elemento `td` dentro del quinto elemento `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("td"));
```

Tu elemento `td` debe tener el contenido de texto `Lunch Break`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.textContent.trim(), "Lunch Break");
``` 

Tu elemento `td` debe tener un atributo `colspan` establecido en `3`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
  </table>
</body>
</html>
```
