---
id: 68c34c608cf9e53f5d36b1d5
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

A continuación, agrega una tercera fila a la tabla. Comienza agregando otro elemento `tr`. Dentro de ese elemento `tr`, agrega un elemento `th` con un atributo `scope` establecido en `"row"` y el contenido de texto `11:00 AM`. Luego, debajo de ese elemento `th`, agrega un elemento `td` con el contenido de texto `Break`. 

# --hints--

Debes tener un tercer elemento `tr` dentro del elemento `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 3);
```

Debes tener un elemento `th` dentro del tercer elemento `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("th"));
```

Tu elemento `th` debe tener el contenido de texto `11:00 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.textContent.trim(), "11:00 AM");
```

Tu elemento `th` debe tener un atributo `scope` establecido en `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.getAttribute("scope"), "row");
```

Debes tener un elemento `td` dentro del tercer elemento `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("td"));
```

Tu elemento `td` debe tener el contenido de texto `Break`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.textContent.trim(), "Break");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

--fcc-editable-region--
      
--fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
