---
id: 68c343692285523aef0fc6bf
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

El atributo `scope` se usa para especificar si una celda de encabezado es un encabezado para una fila, columna o grupo de filas o columnas. Aquí tienes un ejemplo:

```html
<th scope="col">Example Header</th>
```

Esto ayuda a los lectores de pantalla a entender la relación entre las celdas de encabezado y de datos.

Para todos los elementos `th`, agrega un atributo `scope` con un valor de `col`.

# --hints--

Cada elemento `th` debe tener un atributo `scope` con el valor `"col"`.

```js
const tableHeaders = document.querySelectorAll("th");
tableHeaders.forEach(header => {
  assert.equal(header?.getAttribute("scope"), "col");
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
      --fcc-editable-region--
        <th>Time</th>
        <th>Track A</th>
        <th>Track B</th>
        <th>Track C</th>
      --fcc-editable-region--
      </tr>
    </thead>
  </table>
</body>
</html>
```
