---
id: 5d5a813321b9e3db6c106a46
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

JavaScript — це потужна мова, яка дозволяє створювати інтерактивні вебсайти.

*Примітка*: Для всіх наступних проєктів у цій навчальній програмі вам знадобиться базовий рівень знань HTML і CSS. Якщо ви новачок у HTML і CSS, будь ласка, пройдіть [Responsive Web Design Certification](https://www.freecodecamp.org/learn/2022/responsive-web-design/).

Щоб почати, створіть стандартний шаблонний код HTML з `DOCTYPE`, `html`, `head` і `body`, потім додайте тег `meta` для `charset`. Додайте елемент `title` і використайте для нього текст `RPG - Dragon Repeller`. Включіть тег `link` для вашого стилю, щоб підключити файл `styles.css`.

Нарешті, створіть елемент `div` з `id`, встановленим у `game`, всередині вашого `body`.

# --hints--

Ваш код має містити референс `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Після референса `DOCTYPE` має бути пробіл.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Визначте тип документа як `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Закрийте оголошення `DOCTYPE` кінцевим тегом `>` після типу.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Ваш елемент `html` має мати початковий тег. Не забудьте атрибут `lang`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ваш елемент `html` має мати кінцевий тег.

```js
assert(code.match(/<\/html\s*>/));
```

Оголошення `DOCTYPE` має бути на початку вашого HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Ви маєте мати початковий тег `head`.

```js
assert(code.match(/<head\s*>/i));
```

Ви маєте мати кінцевий тег `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Ви маєте мати початковий тег `body`.

```js
assert(code.match(/<body\s*>/i));
```

Ви маєте мати кінцевий тег `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Елементи `head` і `body` мають бути братськими.

```js
assert.match(code, /<head\s*>.*<\/head\s*>.*<body\s*>.*<\/body\s*>/s)
```

Елемент `head` має бути всередині елемента `html`.

```js
assert.match(code, /<html[^>]*>.*<head\s*>.*<\/head\s*>.*<\/html\s*>/s);
```

Елемент `body` має бути всередині елемента `html`.

```js
assert.match(code, /<html[^>]*>.*<body\s*>.*<\/body\s*>.*<\/html\s*>/s);
```

Ваш код має містити елемент `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Ваш елемент `meta` має мати атрибут `charset` зі значенням `UTF-8`.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

Ваш код має містити елемент `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Ви маєте мати кінцевий тег `title`.

```js
assert(code.match(/<\/title\s*>/i));
```

Ваш код має містити елемент `link`.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

Ви маєте мати елемент `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Ваш елемент `div` має мати атрибут `id` зі значенням `game`.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

Ваш елемент `div` має бути всередині елемента `body`.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
