---
id: 62a1166ed9a56d439c0770e7
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Створіть три `span` елементи всередині вашого `#stats` елемента. Додайте кожному з них клас `stat`. Потім першому надайте текст `XP: 0`, другому — текст `Health: 100`, а третьому — текст `Gold: 50`.

# --hints--

У вас має бути три `span` елементи всередині `#stats` елемента.

```js
const spans = document.querySelectorAll('#stats > span');
assert.equal(spans?.length, 3);
```

Новим трьом `span` елементам слід додати клас `stat`.

```js
assert.lengthOf(document.querySelectorAll('#stats > .stat'), 3);
```

Перший `.stat` елемент має містити наданий текст `XP: 0`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[0]?.innerText === 'XP: 0');
```

Другий `.stat` елемент має містити наданий текст `Health: 100`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[1]?.innerText === 'Health: 100');
```

Третій `.stat` елемент має містити наданий текст `Gold: 50`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[2]?.innerText === 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
      <div id="stats">

      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
