---
id: 62a115879a6d51422652cbfc
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

У вашій рольовій грі користувачі зможуть відстежувати свої характеристики, купувати зброю та боротися з монстрами. Перш ніж продовжити з інтерактивною частиною гри на JavaScript, спочатку потрібно створити HTML-елементи, які відображатимуть інформацію про гру.

Створіть чотири `div` елементи всередині вашого `#game` елемента. Присвойте їм відповідні `id` значення в такому порядку: `stats`, `controls`, `monsterStats` та `text`.

# --hints--

Вам потрібно створити чотири нові `div` елементи.

```js
assert.equal(document.querySelectorAll('div')?.length, 5);
```

Одному з нових `div` елементів потрібно присвоїти `id` зі значенням `stats`.

```js
assert.exists(document.querySelector('div#stats'));
```

Одному з нових `div` елементів потрібно присвоїти `id` зі значенням `controls`.

```js
assert.exists(document.querySelector('div#controls'));
```

Одному з нових `div` елементів потрібно присвоїти `id` зі значенням `monsterStats`.

```js
assert.exists(document.querySelector('div#monsterStats'));
```

Одному з нових `div` елементів потрібно присвоїти `id` зі значенням `text`.

```js
assert.exists(document.querySelector('div#text'));
```

Розмістіть нові `div` елементи у правильному порядку.

```js
function __t(a, b) {
  return document.querySelector(a)?.nextElementSibling?.getAttribute('id') === b;
}
assert(__t('div#stats','controls') && __t('div#controls','monsterStats') && __t('div#monsterStats','text'));
```

Розмістіть нові `div` елементи всередині `#game` елемента.

```js
assert.equal(document.querySelector('#game')?.children?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
