---
id: 62a240c67f3dbb1a1e6d95ee
title: Крок 23
challengeType: 0
dashedName: step-23
---

# --description--

Задайте вашому `#game` максимальну ширину `500px` і максимальну висоту `400px`. Встановіть `background-color` у `#ffffff`, а `color` у `#ffffff`. 

Використайте поля, щоб центровано розмістити елемент, задавши верхнє поле `30px`, нижнє поле `0px`, а ліве і праве поле — `auto`. 

Нарешті, додайте `10px` відступу з усіх чотирьох сторін.

# --hints--

У вас має бути селектор `#game`.

```js
const game = new __helpers.CSSHelp(document).getStyle('#game');
assert.exists(game);
```

Ваш селектор `#game` має мати поле `max-width` зі значенням `500px`.

```js
const maxWidth = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('max-width');
assert.equal(maxWidth, '500px');
```

Ваш селектор `#game` має мати поле `max-height` зі значенням `400px`.

```js
const maxHeight = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('max-height');
assert.equal(maxHeight, '400px');
```

Ваш селектор `#game` має мати поле `background-color` зі значенням `#ffffff`.

```js
const background = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('background-color');
assert.equal(background, 'rgb(255, 255, 255)');
```

Ваш селектор `#game` має мати поле `color` зі значенням `#ffffff`.

```js
const color = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('color');
assert.equal(color, 'rgb(255, 255, 255)');
```

Ваш селектор `#game` має мати відступ `margin-top`, встановлений у `30px`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-top');
assert.equal(margin, '30px');
```

Ваш селектор `#game` має мати відступ `margin-left`, встановлений у `auto`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-left');
assert.equal(margin, 'auto');
```

Ваш селектор `#game` має мати відступ `margin-right`, встановлений у `auto`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-right');
assert.equal(margin, 'auto');
```


Ваш селектор `#game` має мати відступ `margin-bottom`, встановлений у `0px`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-bottom');
assert.equal(margin, '0px');
```

Ваш селектор `#game` має мати відступ `10px` з усіх сторін.

```js
const padding = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('padding');
assert.equal(padding, '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
