---
id: 62a257659d0d1e2456f24ba2
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Перед тим як почати писати код свого проєкту, слід перемістити його у власний файл, щоб усе було організовано.

Видаліть свій рядок `console.log("Hello World");`. Потім додайте порожньому елементу `script` атрибут `src` зі значенням `./script.js`.

# --hints--

У вашому коді не повинно бути рядка `console.log("Hello World");`.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Ваш елемент `script` має мати атрибут `src` зі значенням `./script.js`.

```js
const script = document.querySelector("script[data-src$='script.js']");
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>
      console.log("Hello World");
    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
