---
id: 62a255dae245b52317da824a
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Один із найпотужніших інструментів — це ваша консоль розробника. Залежно від браузера, її можна відкрити, натиснувши `F12` або `Ctrl+Shift+I`. На Mac можна натиснути `Option + ⌘ + C` і вибрати «Console». Також можна натиснути кнопку «Console» над вікном попереднього перегляду, щоб побачити нашу вбудовану консоль.

Консоль розробника показує помилки, які виникають у вашому коді, але ви також можете використовувати її, щоб бачити значення змінних у коді, що дуже допомагає для налагодження.

Додайте рядок `console.log("Hello World");` між вашими тегами `script`. Потім натисніть кнопку «Console», щоб відкрити консоль. Ви повинні побачити текст `"Hello World"`.


# --hints--

Ви повинні додати рядок `console.log("Hello World");` у свій код.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;?/);
```

Ваш рядок `console.log("Hello World");` має бути між вашими тегами `script`.

```js
assert.match(code, /<script\s*>\s*console\.log\(\s*('|")Hello World\1\s*\)\s*;?\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
