---
id: 62a3a75d8466a12e009eff76
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Ініціалізуйте ще одну змінну з назвою `health` зі значенням `100`, а також змінну з назвою `gold` зі значенням `50`.

# --hints--

Ви повинні використати `let` для оголошення змінної з назвою `health`.

```js
assert.match(code, /let\s+health/);
```

Ви повинні ініціалізувати змінну з назвою `health` зі значенням `100`.

```js
assert.match(code, /let\s+health\s*=\s*100/);
```

Ви повинні використати `let` для оголошення змінної з назвою `gold`.

```js
assert.match(code, /let\s+gold/);
```

Ви повинні ініціалізувати змінну з назвою `gold` зі значенням `50`.

```js
assert.match(code, /let\s+gold\s*=\s*50/);
```

`health` має мати значення `100`.

```js
assert.equal(health, 100);
```

`gold` має мати значення `50`.

```js
assert.equal(gold, 50);
```

`xp` має й надалі мати значення `0`.

```js
assert.equal(xp, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
--fcc-editable-region--
```
