---
id: 62a3cfc8328d3351b95d4f61
title: Крок 38
challengeType: 0
dashedName: step-38
---

# --description--

Властивість `innerText` керує текстом, який відображається в HTML-елементі. Наприклад:

```html
<p id="info">Demo content</p> 
```

```js
const info = document.querySelector("#info"); 
info.innerText = "Hello World"; 
```

Приклад вище змінить текст елемента `p` з `Demo content` на `Hello World`.

Коли гравець натисне на вашу кнопку `Go to store`, ви хочете змінити кнопки та текст. Видаліть код всередині функції `goStore` і додайте рядок, який оновить текст `button1`, щоб він показував `"Buy 10 health (10 gold)"`.

# --hints--

У вашому коді не повинно бути рядка `console.log("Going to store.");`.

```js
assert.notMatch(code, /console\.log\(\s*('|")Going to store\.\1\s*\)/);
```

Використовуйте точкову нотацію, щоб отримати доступ до властивості `innerText` об’єкта `button1`.

```js
assert.match(code, /button1\.innerText/);
```

Не використовуйте `let` або `const` для доступу до властивості `innerText` об’єкта `button1`.

```js
assert.notMatch(code, /(let|const)\s+button1\.innerText/);
```

Оновіть властивість `innerText` об’єкта `button1`, щоб вона дорівнювала `Buy 10 health (10 gold)`.

```js
assert.match(code, /button1\.innerText\s*=\s*('|")Buy 10 health \(10 gold\)\1/);
```

Оновіть властивість `innerText` у вашій функції `goStore`.

```js
assert.match(goStore.toString(), /button1\.innerText\s*=\s*('|")Buy 10 health \(10 gold\)\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

--fcc-editable-region--
function goStore() {
  console.log("Going to store.");
}
--fcc-editable-region--

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}
```
