---
id: 62a3a0a3c0a4b32915d26a6e
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Видаліть свій рядок `console.log("Hello World");`, щоб почати писати код проєкту.

У попередньому проєкті ви навчились працювати зі змінними та ключовим словом `let` так:

```js
let camperbot;
```

Ви також навчились ініціалізувати змінну значенням так:

```js
let age = 32;
```

Використайте ключове слово `let`, щоб оголосити змінну з ім’ям `xp` і призначити їй значення `0`, яке є числом.

# --hints--

У вашому коді не має бути рядка `console.log("Hello World");`.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Ви маєте використати ключове слово `let` для оголошення змінної.

```js
assert.match(code, /let/);
```

Оголосіть змінну з ім’ям `xp`.

```js
assert.match(code, /let\s+xp/);
```

`xp` має мати значення `0`.

```js
assert.equal(xp, 0);
```

Ви маєте ініціалізувати змінну `xp` значенням `0`. 

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
