---
id: 62a3b79d520a7f3d0e25afd6
title: Крок 15
challengeType: 0
dashedName: step-15
---

# --description--

JavaScript взаємодіє з HTML за допомогою <dfn>Document Object Model</dfn>, або DOM. DOM — це дерево об’єктів, яке представляє HTML. Ви можете отримати доступ до HTML за допомогою об’єкта `document`, який представляє весь ваш HTML-документ.

Один із методів для пошуку конкретних елементів у вашому HTML — це метод `querySelector()`. Метод `querySelector()` приймає CSS-селектор як аргумент і повертає перший елемент, який відповідає цьому селектору. Наприклад, щоб знайти елемент `<h1>` у вашому HTML, ви б написали:

```js
let h1 = document.querySelector("h1");
```

Зверніть увагу, що `h1` — це рядок і відповідає CSS-селектору, який ви б використали.

Створіть змінну `button1` і використайте `querySelector()`, щоб призначити їй ваш елемент з ідентифікатором `id` `button1`. Пам’ятайте, що CSS-селектори `id` мають префікс `#`.

# --hints--

Ви повинні використати `let`, щоб оголосити змінну `button1`.

```js
assert.match(code, /let\s+button1/);
```

Ви повинні використати `document.querySelector()`.

```js
assert.match(code, /document\.querySelector/);
```

Ви повинні використати селектор `"#button1"`.

```js
assert.match(code, /querySelector\(\s*('|")#button1\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--

--fcc-editable-region--
```
