---
id: 62a3b41c9494f937560640ab
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Оголосіть ще дві змінні з іменами `monsterHealth` та `inventory`.

Для вашої змінної `inventory` призначте значення масиву, що містить рядок `"stick"`.

Пам’ятайте, що ви працювали з масивами у попередньому проєкті так:

```js
let exampleArray = ["first", "second", "third"];
```

# --hints--

Ви повинні використати `let` для оголошення вашої змінної `monsterHealth`.

```js
assert.match(code, /let\s+monsterHealth/i);
```

Використовуйте верблюжий регістр для імені вашої змінної `monsterHealth`.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` не повинна мати значення.

```js
assert.isUndefined(monsterHealth);
```

Ви не повинні призначати значення вашій змінній `monsterHealth`.

```js
assert.match(code, /let\s+monsterHealth\s*;?/);
```

Ви повинні використати `let` для оголошення вашої змінної `inventory`.

```js
assert.match(code, /let\s+inventory/i);
```

Ваша змінна `inventory` має залишатися масивом.

```js
assert.isArray(inventory);
```

Ваша змінна `inventory` має містити лише одне значення.

```js
assert.lengthOf(inventory, 1);
```

Ваша змінна `inventory` має містити рядок `"stick"`.

```js
assert.include(inventory, "stick");
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
--fcc-editable-region--
```
