---
id: 62a39f5a5790eb27c1e5d4bf
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Ваш перегляд переключено на ваш новий файл `script.js`. Пам’ятайте, що ви можете використовувати вкладки зверху, щоб перемикатися між файлами.

Додайте свій рядок `console.log("Hello World");` до цього файлу і побачте, як він з’явиться у вашій консолі.

# --hints--

У вашому коді має бути рядок `console.log("Hello World");`.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
