---
id: 62b46e3a8d4be31be5af793d
title: Крок 16
challengeType: 0
dashedName: step-16
---

# --description--

У нас виникла невелика проблема. Ви намагаєтеся зробити запит на сторінці для кнопкового елемента, але ваш `script` тег знаходиться всередині `head` вашого HTML. Це означає, що ваш код виконується раніше, ніж браузер закінчить читати HTML, і ваш `document.querySelector()` не побачить кнопку — бо браузер ще не обробив її.

Щоб виправити це, перемістіть ваш `script` елемент з `head` елемента і розмістіть його в кінці вашого `body` елемента (безпосередньо перед закриваючим тегом `</body>`).

# --hints--

Ваш `script` елемент не повинен бути всередині `head` елемента.

```js
const script = document.querySelector('script[data-src$="script.js"]');
assert.notEqual(script.parentElement.tagName, "HEAD");
```

Ваш `script` елемент має йти в кінці `body` елемента.

```js
const script = document.querySelector('script[data-src$="script.js"]');
const lastChild = document.querySelector('body').lastElementChild;
assert.equal(script, lastChild);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>

  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

let button1 = document.querySelector("#button1");
```
