---
id: 587d8259367417b2b2512c86
title: "implementare l'ordinamento per inserimento"
challengeType: 26
dashedName: implement-insertion-sort
---

# --description--

In questo laboratorio implementerai l'algoritmo di ordinamento per inserimento. Questo metodo funziona costruendo un array ordinato all'inizio della lista. Inizia l'array ordinato con il primo elemento. Poi esamina il successivo elemento e lo scambia all'indietro nell'array ordinato finché non è nella posizione corretta. Continua a iterare attraverso la lista e a scambiare i nuovi elementi all'indietro nella parte ordinata finché non raggiunge la fine. Questo algoritmo ha complessità temporale quadratica nei casi medi e peggiori.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti creare una funzione `insertionSort`.
1. La funzione `insertionSort` dovrebbe prendere un array di numeri interi e restituire un array con gli stessi numeri ordinati dal più piccolo al più grande.

# --hints--

`insertionSort` dovrebbe essere una funzione.

```js
assert.isFunction(insertionSort);
```

`insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` dovrebbe restituire un array che non cambia tranne per l'ordine.

```js
assert.sameMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`insertionSort` dovrebbe restituire un array ordinato (dal più piccolo al più grande).

```js

assert.sameOrderedMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`insertionSort([5, 4, 33, 2, 8])` dovrebbe restituire `[2, 4, 5, 8, 33]`.

```js
assert.sameOrderedMembers(insertionSort([5, 4, 33, 2, 8]), [2, 4, 5, 8, 33])
```

`insertionSort` non dovrebbe usare il metodo integrato `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    insertionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function insertionSort (array) {
  for (let currentIndex = 0; currentIndex < array.length; currentIndex++) {
    let current = array[currentIndex];
    let j = currentIndex - 1;
    while (j > -1 && array[j] > current) {
      array[j + 1] = array[j];
      j--;
    }
    array[j + 1] = current;
  }
  return array;
}
```
