---
id: 63b60a140bf5a321d50a7315
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

사용자가 칼로리 수치를 입력할 식사 유형을 선택하려 할 때, 드롭다운 메뉴와 식사 유형을 추가하는 버튼이 제공되어야 합니다.

먼저 `div` 요소를 만들고 `class` 속성을 `controls` 값으로 할당하세요. 그런 다음 이 `span` 안에 `div` 요소를 중첩하세요.

# --hints--

`div` 안에 새로운 `form` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector('form > div'));
```

`div` 요소는 `fieldset` 요소들 뒤에 와야 합니다.

```js
assert.equal(document.querySelector('form > div')?.previousElementSibling?.tagName, "FIELDSET");
assert.notExists(document.querySelector('form > div')?.nextElementSibling);
```

`div` 요소는 `class` 속성을 `controls` 값으로 설정해야 합니다.

```js
assert.equal(document.querySelector('form > div')?.className, "controls");
```

`span` 안에 `div` 요소를 중첩해야 합니다.

```js
assert.exists(document.querySelector('form > div > span'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
