---
id: 63b608ebf40c871d960fc004
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

`fieldset` 코드 블록이 두 개 더 필요합니다 – 하나는 `snacks` 용, 다른 하나는 `exercise` 용입니다.

# --hints--

`fieldset` 안에 네 번째 `form` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelectorAll('form fieldset')[3]);
```

네 번째 `fieldset` 요소는 세 번째 `fieldset` 요소 뒤에 와야 합니다.

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.previousElementSibling?.tagName, "FIELDSET");
```

네 번째 `fieldset` 요소는 `id`가 `snacks`로 설정되어 있어야 합니다.

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.id, "snacks");
```

네 번째 `fieldset` 요소는 `legend` 요소를 포함해야 합니다.

```js
assert.exists(document.querySelectorAll('form fieldset')[3]?.querySelector('legend'));
```

네 번째 `legend` 요소는 `Snacks` 텍스트를 가져야 합니다.

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.querySelector('legend')?.innerText, "Snacks");
```

네 번째 `fieldset` 요소는 `div` 요소를 포함해야 합니다.

```js
assert.exists(document.querySelectorAll('form fieldset')[3]?.querySelector('div'));
```

네 번째 `div` 요소는 `class`가 `input-container`로 설정되어 있어야 합니다.

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.querySelector('div')?.className, "input-container");
```

`fieldset` 안에 다섯 번째 `form` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelectorAll('form fieldset')[4]);
```

다섯 번째 `fieldset` 요소는 네 번째 `fieldset` 요소 뒤에 와야 합니다.

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.previousElementSibling?.tagName, "FIELDSET");
```

다섯 번째 `fieldset` 요소는 `id`가 `exercise`로 설정되어 있어야 합니다.

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.id, "exercise");
```

다섯 번째 `fieldset` 요소는 `legend` 요소를 포함해야 합니다.

```js
assert.exists(document.querySelectorAll('form fieldset')[4]?.querySelector('legend'));
```

다섯 번째 `legend` 요소는 `Exercise` 텍스트를 가져야 합니다.

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.querySelector('legend')?.innerText, "Exercise");
```

다섯 번째 `fieldset` 요소는 `div` 요소를 포함해야 합니다.

```js
assert.exists(document.querySelectorAll('form fieldset')[4]?.querySelector('div'));
```

다섯 번째 `div` 요소는 `class`가 `input-container`로 설정되어 있어야 합니다.

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.querySelector('div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
