---
id: 63bf511b85b6082e54dc1573
title: 24단계
challengeType: 0
dashedName: step-24
---

# --description--

JavaScript는 문자열의 문자를 다른 문자열로 교체할 수 있는 `.replace()` 메서드를 제공합니다. 이 메서드는 두 개의 인수를 받습니다. 첫 번째 인수는 교체할 문자 수열로, 문자열이나 정규식 패턴일 수 있습니다. 두 번째 인수는 일치하는 수열을 교체할 문자열입니다.

문자열은 불변이므로 `replace` 메서드는 교체된 문자가 반영된 새 문자열을 반환합니다.

이 예제에서는 `replace` 메서드를 사용해 문자열 `l`에서 문자 `1`를 숫자 `hello`로 모두 교체합니다.

```js
"hello".replace(/l/g, "1");
```

`regex`를 사용해 `+`에서 `-`, `str`, 그리고 공백을 모두 빈 문자열로 교체하세요. 이 값을 반환하세요.

# --hints--

`cleanInputString`는 `replace`의 `str` 메서드를 호출해야 합니다.

```js
assert.match(cleanInputString.toString(), /str\.replace\(/);
```

`regex`의 첫 번째 인수로 `replace`을 전달해야 합니다.

```js
assert.match(cleanInputString.toString(), /str\.replace\(\s*regex\s*/);
```

`""`의 두 번째 인수로 `replace`를 전달해야 합니다.

```js
assert.match(cleanInputString.toString(), /str\.replace\(\s*regex\s*,\s*("|')\1\s*\)/);
```

`cleanInputString` 함수는 `replace` 메서드의 결과를 바로 반환해야 합니다.

```js
assert.match(cleanInputString.toString(), /return\s+str\.replace\(\s*regex\s*,\s*("|')\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
          <div class="controls">
            <span>
              <label for="entry-dropdown">Add food or exercise:</label>
              <select id="entry-dropdown" name="options">
                <option value="breakfast" selected>Breakfast</option>
                <option value="lunch">Lunch</option>
                <option value="dinner">Dinner</option>
                <option value="snacks">Snacks</option>
                <option value="exercise">Exercise</option>
              </select>
              <button type="button" id="add-entry">Add Entry</button>
            </span>
          </div>
          <div>
            <button type="submit">
              Calculate Remaining Calories
            </button>
            <button type="button" id="clear">Clear</button>
          </div>
        </form>
        <div id="output" class="output hide"></div>
      </div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js
const calorieCounter = document.getElementById('calorie-counter');
const budgetNumberInput = document.getElementById('budget');
const entryDropdown = document.getElementById('entry-dropdown');
const addEntryButton = document.getElementById('add-entry');
const clearButton = document.getElementById('clear');
const output = document.getElementById('output');
let isError = false;

--fcc-editable-region--
function cleanInputString(str) {
  const regex = /[+-\s]/g;

}
--fcc-editable-region--
```
