---
id: 66b102ee0bdbad34a9f85ba0
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, vas a continuar aprendiendo sobre cadenas creando un Chatbot de Profesor. 

Para comenzar, añade una declaración de `console`, con el mensaje de `"Hi there!"`.

# --hints--

Deberías tener una declaración de `console`.

```js
assert.match(code, /console\.log(.*)/);
```

La declaración de tu consola debería generar el mensaje de `"Hi there!"`.

```js
assert.match(code, /console\.log\((['"])(Hi\s+there!)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
