---
id: 66b102ee0bdbad34a9f85ba0
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, vas a continuar aprendiendo sobre cadenas creando un Chatbot de Profesor.

Para comenzar, añade una declaración de `console`, con el mensaje de `"Hi there!"`.

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

Deberías tener una declaración de `console`.

```js
assert.isAbove(logSpy.calls.length, 0);
```

La declaración de tu consola debería generar el mensaje de `"Hi there!"`.

```js
assert.deepInclude(logSpy.calls, ["Hi there!"]);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
