---
id: 66b59be6ab830800c4df9146
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

Ahora, es momento de usar las variables que creaste en el paso anterior.

In the lectures, you learned how to work with template literals like this:

```js
const name = "John";
const age = 30;

// My name is John and I am 30 years old.
`My name is ${name} and I am ${age} years old.`;
```

Empieza por crear una variable `sentence`.

Usando la sintaxis de literales de plantilla, asigna la cadena de texto `Today, you will learn about [topic variable goes here] in [subject variable goes here].` a la variable `sentence`.

Reemplazarás los marcadores `[topic variable goes here]` y `[subject variable goes here]` con las variables `topic` y `subject` usando la sintaxis `${}`.

Finalmente, registra la variable `sentence` en la consola.

# --hints--

Deberías tener una variable `sentence`.

```js
assert.isNotNull(sentence);
```

Tu variable `sentence` debería ser una cadena de texto.

```js
assert.isString(sentence);
```

Deberías usar la sintaxis de literales de plantilla para asignar la cadena de texto `Today, you will learn about [topic variable goes here] in [subject variable goes here].` a la variable `sentence`. Consulta el ejemplo si necesitas ayuda.

```js
assert.match(code, /Today,\s*you\s+will\s+learn\s+about\s+\$\{topic\}\s+in\s+\$\{subject\}\./);
```

Deberías registrar la variable `sentence` en la consola.

```js
assert.match(code,/console\.log\(\s*sentence\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

--fcc-editable-region--

--fcc-editable-region--
```
