---
id: 66b6e62423e8a031d6c1f03d
title: Paso 8
challengeType: 1
dashedName: step-8
---

# --description--

Ahora es el momento de obtener la longitud de la cadena `topic`.

Puedes usar literales de plantilla dentro de las sentencias de `console` como esta:

```js
const developer = "Jessica";
console.log(`Hello, my name is ${developer}.`);
```

Comienza mostrando el mensaje `Here is an example of using the length property on the word [topic].` en la consola. 

Recuerda reemplazar `[topic]` con la variable `topic`, y usar la sintaxis correcta de literales de plantilla como hiciste en los pasos anteriores.

Luego, agrega una segunda sentencia `console.log` que muestre la longitud de la cadena `topic` en la consola.

# --hints--

Debes usar la sintaxis de literales de plantilla para mostrar el mensaje `Here is an example of using the length property on the word [topic].` en la `console`.

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{topic\}\.`\)/);
```

Debes tener una segunda sentencia `console` para mostrar la longitud de la variable `topic` en la `console`.

```js
assert.match(code, /console\.log\(\s*topic\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);
console.log(subject.length);

--fcc-editable-region--

--fcc-editable-region--
```
