---
id: 66b6ffb42b88e33943788abf
title: Paso 14
challengeType: 1
dashedName: step-14
---

# --description--

Para la última parte del taller, revisará cómo encontrar la posición del índice de una subcadena en una cadena. Recuerde que una subcadena es una parte de una cadena.

Comience creando una variable llamada `learningIsFunSentence` y asígnale el valor de cadena de `"Learning is fun."`.

# --hints--

Debería tener una variable llamada `learningIsFunSentence`.

```js
assert.isNotNull(learningIsFunSentence);
```

Su `learningIsFunSentence` debería tener el valor de una cadena.

```js
assert.isString(learningIsFunSentence);
```

Debería asignar el valor de `"Learning is fun."` a la variable `learningIsFunSentence`.

```js
assert.equal(learningIsFunSentence, "Learning is fun.");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

--fcc-editable-region--

--fcc-editable-region--
```
