---
id: 66b6f586767a1534f3097353
title: Paso 11
challengeType: 1
dashedName: step-11
---

# --description--

Now it is time to access the second letter of the `subject` variable.

Comienza añadiendo una instrucción de `console` que muestre el mensaje `Aquí hay un ejemplo de cómo acceder a la segunda letra en la palabra [subject].` Recuerda reemplazar `[subject]` con el valor actual de la variable `subject` y usar la sintaxis correcta de literales de plantilla.

Luego, añade otra instrucción de `console` que muestre la segunda letra de la variable `subject` usando notación de corchetes y el número de índice correcto.

# --hints--

Tu instrucción `console` debería mostrar el mensaje `Aquí hay un ejemplo de cómo acceder a la segunda letra en la palabra [subject].`. 

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+second\s+letter\s+in\s+the\s+word\s+\$\{subject\}\.`\);?/);
```

Deberías tener otra instrucción de `console` que muestre la segunda letra de la variable `subject` utilizando el índice de notación de corchetes `1`.

```js
assert.match(code, /console\.log\(subject\[1\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

--fcc-editable-region--

--fcc-editable-region--
```
