---
id: 66b6efddeca35833cd6f0b03
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

En lecciones anteriores, aprendiste cómo acceder a caracteres en una cadena así:

```js
const firstName = "Jessica";
// returns "J"
firstName[0];
```  

Recuerda que los números de índice comienzan en `0`, por lo que la primera letra en una cadena siempre estará en el índice `0`.

Comienza agregando otra declaración `console`. 

Dentro de la declaración `console`, muestra la primera letra de la variable `subject` utilizando la notación de corchetes y el número de índice correcto.

# --hints--

Tu declaración `console` debe mostrar la primera letra de la variable `subject` utilizando la notación de corchetes e índice `0`.

```js
assert.match(code, /console\.log\(subject\[0\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
