---
id: 66b70b8a611cbf3bcc5c6c5f
title: Paso 16
challengeType: 1
dashedName: step-16
---

# --description--

En lecciones anteriores, aprendiste cómo trabajar con el método `indexOf` así:

```js
const sentence = "I love to learn.";

// returns index 2
console.log(sentence.indexOf("love"));

// returns -1
console.log(sentence.indexOf("hate"));
```

Recuerda que el método `indexOf` devuelve la posición del índice de la primera aparición de una subcadena en una cadena. Si la subcadena no se encuentra, devuelve `-1`.

Agrega una nueva declaración de `console` que muestre el resultado de usar el método `indexOf` en la variable `learningIsFunSentence` para encontrar la posición del índice de la subcadena `"Learning"`.

# --hints--

Tu `console` debe usar el método `indexOf` en la variable `learningIsFunSentence` para encontrar la posición del índice de la subcadena `"Learning"`.

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(Learning)\1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";

console.log("Here are examples of finding the positions of substrings in the sentence.");

--fcc-editable-region--

--fcc-editable-region--
```
