---
id: 690ddb5e6ee94eb73ed172b0
title: Eseguire il debug di un modulo di donazione
challengeType: 25
dashedName: debug-donation-form
# Purposefully removed the demo type because this is a debugging project.
---

# --description--

Un'organizzazione benefica locale ha creato un sito web con un modulo di donazione, ma ci sono diversi problemi da risolvere. Il modulo non è accessibile e presenta alcuni errori di sintassi HTML.

Il tuo compito è correggere tutti gli errori affinché il modulo funzioni correttamente e sia accessibile a tutti gli utenti. Completa gli elementi nelle storie utente qui sotto e clicca su "Esegui i test" per verificare se hai risolto tutti gli errori.

**Storie utente:**

1. Gli elementi `input` sono elementi void e non dovrebbero avere tag di chiusura. Rimuovi tutti i tag di chiusura `</input>` dal modulo.
2. Aggiungi elementi `label` per ogni campo di input del modulo in modo che gli utenti sappiano a cosa serve ciascun campo. Il testo dell'etichetta dovrebbe corrispondere a quello attualmente accanto a ogni input.
3. Il tipo di input `Email Address:` dovrebbe essere un `email` invece di `text`.
4. Dovresti associare ogni elemento `label` con il corrispondente elemento `input` usando l'attributo `for` sul `label` e un attributo `id` corrispondente sul `input`.
5. Aggiungi l'attributo `required` ai campi di input di tipo testo, email e numero (ma non alla casella di spunta o al pulsante di invio) per assicurarti che gli utenti compilino le informazioni richieste.

# --hints--

Non dovresti avere alcun tag di chiusura `</input>` nel tuo codice.

```js
assert.notMatch(code, /<\/input>/);
```

Dovresti avere esattamente cinque elementi `input` nel tuo modulo.

```js
assert.lengthOf(document.querySelectorAll("input"), 5);
```

Dovresti avere esattamente quattro elementi `label` nel tuo modulo.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

Il tuo primo `label` dovrebbe avere il testo `Full Name:`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[0]?.textContent, /Full\s*Name\s*:/i);
```

Il tuo primo `input` dovrebbe avere un attributo `required`.

```js
const firstInput = document.querySelector("input[name='name']");
assert.isTrue(firstInput?.hasAttribute("required"));
```

Il tuo secondo `label` dovrebbe avere il testo `Email Address:`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[1]?.textContent, /Email\s*Address\s*:/i);
```

Il tuo secondo `input` dovrebbe avere un attributo `required`.

```js
const secondInput = document.querySelector("input[name='email']");
assert.isTrue(secondInput?.hasAttribute("required"));
```

Il tuo terzo `label` dovrebbe avere il testo `Donation Amount ($):`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[2]?.textContent, /Donation\s*Amount\s*\(\$\)\s*:/i);
```

Il tuo terzo `input` dovrebbe avere un attributo `required`.

```js
const thirdInput = document.querySelector("input[name='amount']");
assert.isTrue(thirdInput?.hasAttribute("required"));
```

Il tuo quarto `label` dovrebbe avere il testo `Subscribe`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[3]?.textContent, /Subscribe/i);
```

Dovresti associare ogni elemento `input` con un elemento `label` usando gli attributi `for` e `id`.

```js 
const els = document.querySelectorAll('input:not([type="submit"])'); 
assert.isNotEmpty(els); 
els.forEach(input => { 
  const label = document.querySelector(`label[for="${input.id}"]`); 
  assert.exists(label);  
  assert.equal(label.tagName, "LABEL"); 
})
``` 

Il tuo `input` per l'indirizzo email dovrebbe essere di tipo `email`.

```js
const emailInput = document.querySelector("input[name='email']");
assert.strictEqual(emailInput?.getAttribute("type"), "email");
```

La tua casella di spunta `input` non dovrebbe avere un attributo `required`.

```js
const checkboxInput = document.querySelector("input[name='newsletter']");
assert.isFalse(checkboxInput?.hasAttribute("required"));
```

Il tuo pulsante di invio `button` non dovrebbe avere un attributo `required`.

```js
const submitButton = document.querySelector("input[type='submit']");
assert.isFalse(submitButton?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    
    Full Name:
    <input type="text" name="name"></input>

    Email Address:
    <input type="text" name="email">

    Donation Amount ($):
    <input type="number" name="amount"></input>

    <input type="checkbox" name="newsletter"></input>
    Subscribe

    <input type="submit" value="Send"></input>
  </form>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    <label for="full-name">Full Name:</label>
    <input id="full-name" type="text" name="name" required>

    <label for="email">Email Address:</label>
    <input id="email" type="email" name="email" required>

    <label for="amount">Donation Amount ($):</label>
    <input id="amount" type="number" name="amount" required>

    <label for="newsletter">Subscribe</label>
    <input id="newsletter" type="checkbox" name="newsletter">

    <input type="submit" value="Send">
  </form>
</body>
</html>
```
