---
id: 68cabc534d2e741f33a3fe1d
title: Einen Spielcharakter-Statistik-Tracker erstellen
challengeType: 27
dashedName: lab-game-character-stats
---

# --description--

In diesem Labor erstellen Sie einen Spielcharakter-Statistik-Tracker. Das Programm ermöglicht es Ihnen, einen Charakter mit bestimmten Attributen zu erstellen, diese Attribute zu aktualisieren und die aktuellen Statistiken des Charakters abzurufen.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Erstellen Sie eine Klasse namens `GameCharacter`, die einen Spielcharakter darstellt und die Charakterstatistiken verwaltet.
2. Wenn ein neues `GameCharacter`-Objekt instanziiert wird, sollte es folgende Attribute haben:
    - `_name` wird auf den String gesetzt, der zum Zeitpunkt der Instanziierung übergeben wird.
    - `_health` wird auf `100` gesetzt.
    - `_mana` wird auf `50` gesetzt.
    - `_level` wird auf `1` gesetzt.
3. Erstellen Sie eine `name`-Eigenschaft für den schreibgeschützten Zugriff auf den Namen des Charakters.
4. Für die `health`-Eigenschaft:
     - Definieren Sie einen Getter, der die aktuelle Gesundheit zurückgibt.
     - Definieren Sie einen Setter, der verhindert, dass die Gesundheit unter `0` gesetzt wird, und die Gesundheit auf `100` begrenzt.
5. Für die `mana`-Eigenschaft:
     - Definieren Sie einen Getter, der das aktuelle Mana zurückgibt.
     - Definieren Sie einen Setter, der verhindert, dass das Mana unter `0` gesetzt wird, und das Mana auf `50` begrenzt.
6. Erstellen Sie einen Getter für `level`, der das aktuelle Level des Charakters zurückgibt.
7. Definieren Sie eine Methode namens `level_up`, die:
    - Das Level des Charakters um 1 erhöht.
    - Die Gesundheit auf `100` und das Mana auf `50` zurücksetzt, indem die entsprechenden Eigenschafts-Setter verwendet werden.
    - Eine Nachricht in der Form `<name> leveled up to <level>!` ausgibt (wobei `<name>` und `<level>` durch den Namen des Charakters bzw. das neue Level ersetzt werden sollten).
8. Definieren Sie eine `__str__`-Methode, die einen formatierten String zurückgibt, der Folgendes enthält:
    - Den Namen des Charakters.
    - Das Level des Charakters.
    - Die aktuelle Gesundheit des Charakters.
    - Das aktuelle Mana des Charakters.
   Zum Beispiel sollte ein Charakter namens `Kratos`, direkt nach der Instanziierung, wie folgt dargestellt werden:

   ```md
   Name: Kratos
   Level: 1
   Health: 100
   Mana: 50
   ```

## Beispiel zur Verwendung

```python
hero = GameCharacter('Kratos') # Creates a new character named Kratos
print(hero)  # Displays the character's stats

hero.health -= 30  # Decreases health by 30
hero.mana -= 10    # Decreases mana by 10
print(hero)  # Displays the updated stats

hero.level_up()  # Levels up the character
print(hero)  # Displays the stats after leveling up
```

# --hints--

Sie sollten eine `GameCharacter`-Klasse haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("GameCharacter")`)) })
```

Sie sollten eine `__init__`-Methode in der `GameCharacter`-Klasse haben.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__init__")`) 
})
```

Die `__init__`-Methode sollte zwei Parameter haben, wobei der erste `self` ist.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.__init__).parameters)
    assert params[0] == "self"
    assert len(params) == 2
    `) 
})
```

Sie sollten `self._name` im `__init__`-Konstruktor auf den zum Zeitpunkt der Instanziierung übergebenen String setzen.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._name == "Kratos"
    `) 
})
```

Sie sollten `100` in der `self._health`-Methode auf `__init__` zuweisen.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._health == 100
    `) 
})
```

Sie sollten `50` in der `self._mana`-Methode auf `__init__` zuweisen.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._mana == 50
    `) 
})
```

Sie sollten `1` in der `self._level`-Methode auf `__init__` zuweisen.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._level == 1
    `) 
})
```

Sie sollten eine `name`-Methode in der `GameCharacter`-Klasse haben.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("name")`) 
})
```

Die `name`-Methode sollte nur einen Parameter haben, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_args("self")`) 
})
```

Sie sollten `self._name` in der `name`-Methode zurückgeben.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_return("self._name")`)
})
```

Die `name`-Methode sollte einen `@property`-Decorator haben.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_decorators("property")`) 
})
```

Sie sollten eine `health`-Methode in der `GameCharacter`-Klasse haben.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("health")`) 
})
```

Die `health`-Methode sollte nur einen Parameter haben, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_args("self")`) 
})
```

Sie sollten `self._health` in der `health`-Methode zurückgeben.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_return("self._health")`)
})
```

Die `health`-Methode sollte einen `@property`-Decorator haben.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_decorators("property")`) 
})
```

Sie sollten eine `health`-Setter-Methode erstellen.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1]`) 
})
```

Die `health`-Setter-Methode sollte einen `@health.setter`-Decorator haben.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1].has_decorators("health.setter")`) 
})
```

Die `health`-Setter-Methode sollte zwei Parameter haben, wobei der erste `self` ist.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.health.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

Sie sollten `self._health` auf `0` setzen, wenn der Wert, der dem `health`-Setter übergeben wird, kleiner als `0` ist.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = -1
    assert char.health == 0
    char.health = -10
    assert char.health == 0
`)})
```

Sie sollten `self._health` auf `100` setzen, wenn der Wert, der dem `health`-Setter übergeben wird, größer als `100` ist.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 101
    assert char.health == 100
    char.health = 200
    assert char.health == 100
`)})
```

Sie sollten `self._health` auf den Wert setzen, der dem `health`-Setter übergeben wird, wenn der Wert zwischen `0` und `100` liegt.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 95
    assert char.health == 95
    char.health = 1
    assert char.health == 1
    char.health = 0
    assert char.health == 0
    char.health = 100
    assert char.health == 100
`)})
```

Sie sollten eine `mana`-Methode in der `GameCharacter`-Klasse haben.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("mana")`) 
})
```

Die `mana`-Methode sollte nur einen Parameter haben, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_args("self")`) 
})
```

Sie sollten `self._mana` in der `mana`-Methode zurückgeben.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_return("self._mana")`)
})
```

Die `mana`-Methode sollte einen `@property`-Decorator haben.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_decorators("property")`) 
})
```

Sie sollten eine `mana`-Setter-Methode erstellen.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1]`) 
})
```

Die `mana`-Setter-Methode sollte einen `@mana.setter`-Decorator haben.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1].has_decorators("mana.setter")`) 
})
```

Die `mana`-Setter-Methode sollte zwei Parameter haben, wobei der erste `self` ist.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.mana.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

Sie sollten `self._mana` auf `0` setzen, wenn der Wert, der dem `mana`-Setter übergeben wird, kleiner als `0` ist.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = -1
    assert char.mana == 0
    char.mana = -10
    assert char.mana == 0
`)})
```

Sie sollten `self._mana` auf `50` setzen, wenn der Wert, der dem `mana`-Setter übergeben wird, größer als `50` ist.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 51
    assert char.mana == 50
    char.mana = 100
    assert char.mana == 50
`)})
```

Sie sollten `self._mana` auf den Wert setzen, der dem `mana`-Setter übergeben wird, wenn der Wert zwischen `0` und `50` liegt.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 25
    assert char.mana == 25
    char.mana = 0
    assert char.mana == 0
    char.mana = 50
    assert char.mana == 50
`)})
```

Sie sollten eine `level`-Methode in der `GameCharacter`-Klasse haben.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level")`) 
})
```

Die `level`-Methode sollte nur einen Parameter haben, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_args("self")`) 
})
```

Sie sollten `self._level` in der `level`-Methode zurückgeben.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_return("self._level")`)
})
```

Die `level`-Methode sollte einen `@property`-Decorator haben.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_decorators("property")`) 
})
```

Sie sollten eine `level_up`-Methode in der `GameCharacter`-Klasse haben.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level_up")`) 
})
```

Die `level_up`-Methode sollte nur einen Parameter haben, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_args("self")`) 
})
```

Sie sollten `self._level` in der `1`-Methode um `level_up` erhöhen.

```js
({ 
  test: () => runPython(`
    node = _Node(_code).find_class("GameCharacter").find_function("level_up")
    assert (
        node.has_stmt("self._level += 1") or
        node.has_stmt("self._level = self._level + 1")
    )
`)
})
```

Sie sollten `self.health` in der `100`-Methode auf `level_up` setzen.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.health = 100")`) 
})
```

Sie sollten `self.mana` in der `50`-Methode auf `level_up` setzen.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.mana = 50")`) 
})
```

Die `level_up`-Methode sollte `<name> leveled up to <level>!` ausgeben (wobei `<name>` und `<level>` durch den Namen des Charakters bzw. das neue Level ersetzt werden sollten).

```js
({
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.level_up()
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
            
    assert "test leveled up to 2!" in output
`)
})
```

Sie sollten eine `__str__`-Methode in der `GameCharacter`-Klasse haben.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__str__")`) 
})
```

Die `__str__`-Methode sollte nur einen Parameter haben, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("__str__").has_args("self")`) 
})
```

Ihre `__str__`-Methode sollte einen String mit den Statistiken des Charakters im vorgegebenen Format zurückgeben.

```js
({ 
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.mana = 15
    gc.health = 77
    print(gc)
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
    expected = """Name: test
    Level: 1
    Health: 77
    Mana: 15"""
    assert expected in output
`)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class GameCharacter:
    def __init__(self, name):
        self._name = name
        self._health = 100
        self._mana = 50
        self._level = 1

    @property
    def name(self):
        return self._name

    @property
    def health(self):
        return self._health

    @health.setter
    def health(self, value):
        if value < 0:
            self._health = 0
        elif value > 100:
            self._health = 100
        else:
            self._health = value

    @property
    def mana(self):
        return self._mana

    @mana.setter
    def mana(self, value):
        if value < 0:
            self._mana = 0
        elif value > 50:
            self._mana = 50
        else:
            self._mana = value

    @property
    def level(self):
        return self._level

    def level_up(self):
        self._level += 1
        self.health = 100
        self.mana = 50
        print(f"{self._name} leveled up to {self._level}!")

    def __str__(self):
        return (
            f"Name: {self._name}\n"
            f"Level: {self._level}\n"
            f"Health: {self._health}\n"
            f"Mana: {self._mana}"
        )
```
