---
id: 56104e9e514f539506016a5c
title: Iterieren ungerader Zahlen mit einer for-Schleife
challengeType: 1
forumTopicId: 18212
dashedName: iterate-odd-numbers-with-a-for-loop
---

# --description--

For-Schleifen müssen nicht immer eins nach dem anderen iterieren. Indem wir unseren `final-expression` ändern, können wir in Schritten von geraden Zahlen zählen.

Wir beginnen bei `i = 0` und iterieren, solange `i < 10` gilt. Wir erhöhen `i` bei jeder Schleife um 2 mit `i += 2`.

```js
const ourArray = [];

for (let i = 0; i < 10; i += 2) {
  ourArray.push(i);
}
```

`ourArray` enthält jetzt `[0, 2, 4, 6, 8]`. Ändern wir unsere `initialization`, damit wir in Schritten von ungeraden Zahlen zählen können.

# --instructions--

Fügen Sie die ungeraden Zahlen von 1 bis 9 mit einer `myArray`-Schleife zu `for` hinzu.

# --hints--

Sie sollten dafür eine `for`-Schleife verwenden.

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

`myArray` sollte gleich `[1, 3, 5, 7, 9]` sein.

```js
assert.deepEqual(myArray, [1, 3, 5, 7, 9]);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 1; i < 10; i += 2) {
  myArray.push(i);
}
```
