---
id: 56105e7b514f539506016a5e
title: Rückwärts zählen mit einer for-Schleife
challengeType: 1
forumTopicId: 16808
dashedName: count-backwards-with-a-for-loop
---

# --description--

Eine for-Schleife kann auch rückwärts zählen, solange wir die richtigen Bedingungen definieren können.

Um bei jeder Iteration um zwei zu dekrementieren, müssen wir unsere Initialisierung, Bedingung und den abschließenden Ausdruck ändern.

Wir starten bei `i = 10` und durchlaufen die Schleife, solange `i > 0` gilt. Wir dekrementieren `i` bei jeder Schleife um 2 mit `i -= 2`.

```js
const ourArray = [];

for (let i = 10; i > 0; i -= 2) {
  ourArray.push(i);
}
```

`ourArray` wird nun `[10, 8, 6, 4, 2]` enthalten. Ändern wir unsere Initialisierung und den abschließenden Ausdruck, damit wir rückwärts in Zweierschritten zählen und ein Array absteigender ungerader Zahlen erstellen können.

# --instructions--

Fügen Sie die ungeraden Zahlen von 9 bis 1 mit einer `myArray`-Schleife zu `for` hinzu.

# --hints--

Sie sollten hierfür eine `for`-Schleife verwenden.

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

Sie sollten die Array-Methode `push` verwenden.

```js
assert(__helpers.removeJSComments(code).match(/myArray.push/));
```

`myArray` sollte gleich `[9, 7, 5, 3, 1]` sein.

```js
assert.deepEqual(myArray, [9, 7, 5, 3, 1]);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 9; i > 0; i -= 2) {
  myArray.push(i);
}
```
