---
id: 56533eb9ac21ba0edf2244b6
title: Escape-Sequenzen in Strings
challengeType: 1
forumTopicId: 17567
dashedName: escape-sequences-in-strings
---

# --description--

Anführungszeichen sind nicht die einzigen Zeichen, die innerhalb eines Strings <dfn>escaped</dfn> werden können. Escape-Sequenzen ermöglichen es Ihnen, Zeichen zu verwenden, die Sie sonst möglicherweise nicht in einem String verwenden könnten.

<table><thead><tr><th>Code</th><th>Ausgabe</th></tr></thead><tbody><tr><td><code>\'</code></td><td>einfaches Anführungszeichen</td></tr><tr><td><code>\"</code></td><td>doppeltes Anführungszeichen</td></tr><tr><td><code>\\</code></td><td>Backslash</td></tr><tr><td><code>\n</code></td><td>Zeilenumbruch</td></tr><tr><td><code>\t</code></td><td>Tabulator</td></tr><tr><td><code>\r</code></td><td>Wagenrücklauf</td></tr><tr><td><code>\b</code></td><td>Rücktaste</td></tr><tr><td><code>\f</code></td><td>Seitenvorschub</td></tr></tbody></table>

*Hinweis: Der Backslash selbst muss escaped werden, damit er als Backslash angezeigt wird.*

# --instructions--

Weisen Sie die folgenden drei Textzeilen mit Escape-Sequenzen der einzelnen Variablen `myStr` zu.

<pre>
FirstLine
    \SecondLine
ThirdLine
</pre>

Sie müssen Escape-Sequenzen verwenden, um Sonderzeichen korrekt einzufügen. Außerdem müssen Sie den Abstand so einhalten, wie er oben aussieht, ohne Leerzeichen zwischen Escape-Sequenzen oder Wörtern.

**Hinweis:** Die Einrückung für `SecondLine` wird mit dem Tabulator-Escape-Zeichen erreicht, nicht mit Leerzeichen.

# --hints--

`myStr` sollte keine Leerzeichen enthalten

```js
assert(!/ /.test(myStr));
```

`myStr` sollte die Strings `FirstLine`, `SecondLine` und `ThirdLine` enthalten (denken Sie an Groß- und Kleinschreibung)

```js
assert(
  /FirstLine/.test(myStr) && /SecondLine/.test(myStr) && /ThirdLine/.test(myStr)
);
```

`FirstLine` sollte von dem Zeilenumbruchzeichen `\n` gefolgt sein

```js
assert(/FirstLine\n/.test(myStr));
```

`myStr` sollte ein Tabulatorzeichen `\t` enthalten, das einem Zeilenumbruchzeichen folgt

```js
assert(/\n\t/.test(myStr));
```

`SecondLine` sollte von dem Backslash-Zeichen `\` vorangestellt sein

```js
assert(/\\SecondLine/.test(myStr));
```

Zwischen `SecondLine` und `ThirdLine` sollte ein Zeilenumbruchzeichen stehen

```js
assert(/SecondLine\nThirdLine/.test(myStr));
```

`myStr` sollte nur die in der Anleitung gezeigten Zeichen enthalten

```js
assert(myStr === 'FirstLine\n\t\\SecondLine\nThirdLine');
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "FirstLine\n\t\\SecondLine\nThirdLine";
```
