---
id: 56533eb9ac21ba0edf2244d9
title: Vergleiche mit dem logischen Oder-Operator
challengeType: 1
forumTopicId: 16800
dashedName: comparisons-with-the-logical-or-operator
---

# --description--

Der <dfn>logische Oder</dfn>-Operator (`||`) gibt `true` zurück, wenn einer der <dfn>Operanden</dfn> `true` ist. Andernfalls gibt er `false` zurück.

Der <dfn>logische Oder</dfn>-Operator besteht aus zwei senkrechten Strichen: (`||`). Diese befinden sich typischerweise zwischen der Rücktaste und der Eingabetaste.

Das folgende Muster sollte Ihnen von vorherigen Wegpunkten vertraut sein.

```js
if (num > 10) {
  return "No";
}
if (num < 5) {
  return "No";
}
return "Yes";
```

Dieser Code gibt `Yes` zurück, wenn `num` zwischen `5` und `10` liegt (einschließlich `5` und `10`). Dieselbe Logik kann mit dem <dfn>logischen Oder</dfn>-Operator geschrieben werden.

```js
if (num > 10 || num < 5) {
  return "No";
}
return "Yes";
```

# --instructions--

Kombinieren Sie die beiden `if`-Anweisungen zu einer einzigen Anweisung, die den String `Outside` zurückgibt, wenn `val` nicht zwischen `10` und `20` liegt, einschließlich beider Grenzen. Andernfalls geben Sie den String `Inside` zurück.

# --hints--

Sie sollten den `||`-Operator einmal verwenden

```js
assert(__helpers.removeJSComments(code).match(/\|\|/g).length === 1);
```

Sie sollten nur eine `if`-Anweisung haben

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`testLogicalOr(0)` sollte den String `Outside` zurückgeben

```js
assert(testLogicalOr(0) === 'Outside');
```

`testLogicalOr(9)` sollte den String `Outside` zurückgeben

```js
assert(testLogicalOr(9) === 'Outside');
```

`testLogicalOr(10)` sollte den String `Inside` zurückgeben

```js
assert(testLogicalOr(10) === 'Inside');
```

`testLogicalOr(15)` sollte den String `Inside` zurückgeben

```js
assert(testLogicalOr(15) === 'Inside');
```

`testLogicalOr(19)` sollte den String `Inside` zurückgeben

```js
assert(testLogicalOr(19) === 'Inside');
```

`testLogicalOr(20)` sollte den String `Inside` zurückgeben

```js
assert(testLogicalOr(20) === 'Inside');
```

`testLogicalOr(21)` sollte den String `Outside` zurückgeben

```js
assert(testLogicalOr(21) === 'Outside');
```

`testLogicalOr(25)` sollte den String `Outside` zurückgeben

```js
assert(testLogicalOr(25) === 'Outside');
```

# --seed--

## --seed-contents--

```js
function testLogicalOr(val) {
  // Only change code below this line

  if (val) {
    return "Outside";
  }

  if (val) {
    return "Outside";
  }

  // Only change code above this line
  return "Inside";
}

testLogicalOr(15);
```

# --solutions--

```js
function testLogicalOr(val) {
  if (val < 10 || val > 20) {
    return "Outside";
  }
  return "Inside";
}
```
