---
id: 56533eb9ac21ba0edf2244b5
title: Literal-Anführungszeichen in Strings maskieren
challengeType: 1
forumTopicId: 17568
dashedName: escaping-literal-quotes-in-strings
---

# --description--

Wenn Sie einen String definieren, müssen Sie ihn mit einfachen oder doppelten Anführungszeichen beginnen und beenden. Was passiert, wenn Sie ein wörtliches Anführungszeichen `"` oder `'` innerhalb Ihres Strings benötigen?

In JavaScript können Sie ein Anführungszeichen <dfn>maskieren</dfn>, damit es nicht als Ende des Strings interpretiert wird, indem Sie einen <dfn>Backslash</dfn> (`\`) vor das Anführungszeichen setzen.

```js
const sampleStr = "Alan said, \"Peter is learning JavaScript\".";
```

Dies signalisiert JavaScript, dass das folgende Anführungszeichen nicht das Ende des Strings ist, sondern stattdessen im String erscheinen soll. Wenn Sie dies also in der Konsole ausgeben würden, bekämen Sie:

```js
Alan said, "Peter is learning JavaScript".
```

# --instructions--

Verwenden Sie <dfn>Backslashes</dfn>, um eine Zeichenkette der Variablen `myStr` zuzuweisen, sodass bei einer Ausgabe in der Konsole Folgendes erscheint:

```js
I am a "double quoted" string inside "double quotes".
```

# --hints--

Sie sollten zwei doppelte Anführungszeichen (`"`) und vier maskierte doppelte Anführungszeichen (`\"`) verwenden.

```js
assert(__helpers.removeJSComments(code).match(/\\"/g).length === 4 && __helpers.removeJSComments(code).match(/[^\\]"/g).length === 2);
```

Die Variable `myStr` sollte den String enthalten: `I am a "double quoted" string inside "double quotes".`

```js
assert(/I am a "double quoted" string inside "double quotes(\."|"\.)$/.test(myStr));
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "I am a \"double quoted\" string inside \"double quotes\".";
```
