---
id: 56533eb9ac21ba0edf2244a9
title: Variablen mit dem Zuweisungsoperator initialisieren
challengeType: 1
forumTopicId: 301171
dashedName: initializing-variables-with-the-assignment-operator
---

# --description--

Es ist üblich, eine Variable in derselben Zeile, in der sie deklariert wird, mit einem Anfangswert zu initialisieren.

```js
var myVar = 0;
```

Erstellt eine neue Variable namens `myVar` und weist ihr einen Anfangswert von `0` zu.

# --instructions--

Definieren Sie eine Variable `a` mit `var` und initialisieren Sie sie mit dem Wert `9`.

# --hints--

Sie sollten `a` mit dem Wert `9` initialisieren.

```js
assert(/var\s+a\s*=\s*9(\s*;?\s*)$/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
var a = 9;
```
