---
id: 56533eb9ac21ba0edf2244b1
title: Zusammengesetzte Zuweisung mit erweiterter Multiplikation
challengeType: 1
forumTopicId: 16662
dashedName: compound-assignment-with-augmented-multiplication
---

# --description--

Der `*=`-Operator multipliziert eine Variable mit einer Zahl.

```js
myVar = myVar * 5;
```

multipliziert `myVar` mit `5`. Dies kann umgeschrieben werden als:

```js
myVar *= 5;
```

# --instructions--

Konvertieren Sie die Zuweisungen für `a`, `b` und `c`, sodass der `*=`-Operator verwendet wird.

# --hints--

`a` sollte gleich `25` sein.

```js
assert(a === 25);
```

`b` sollte gleich `36` sein.

```js
assert(b === 36);
```

`c` sollte gleich `46` sein.

```js
assert(c === 46);
```

Sie sollten für jede Variable den `*=`-Operator verwenden.

```js
assert(__helpers.removeJSComments(code).match(/\*=/g).length === 3);
```

Sie dürfen den Code oberhalb des angegebenen Kommentars nicht verändern.

```js
assert(
  /let a = 5;/.test(__helpers.removeJSComments(code)) &&
    /let b = 12;/.test(__helpers.removeJSComments(code)) &&
    /let c = 4\.6;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 5;
let b = 12;
let c = 4.6;

// Only change code below this line
a = a * 5;
b = 3 * b;
c = c * 10;
```

# --solutions--

```js
let a = 5;
let b = 12;
let c = 4.6;

a *= 5;
b *= 3;
c *= 10;
```
