---
id: 56533eb9ac21ba0edf2244ab
title: Groß- und Kleinschreibung bei Variablen verstehen
challengeType: 1
forumTopicId: 18334
dashedName: understanding-case-sensitivity-in-variables
---

# --description--

In JavaScript sind alle Variablen- und Funktionsnamen case-sensitive. Das bedeutet, dass die Groß- und Kleinschreibung eine Rolle spielt.

`MYVAR` ist nicht dasselbe wie `MyVar` noch `myvar`. Es ist möglich, mehrere unterschiedliche Variablen mit demselben Namen, aber unterschiedlicher Groß- und Kleinschreibung zu haben. Es wird dringend empfohlen, diese Sprachfunktion aus Gründen der Übersichtlichkeit *nicht* zu verwenden.

**Bewährte Vorgehensweise**

Schreiben Sie Variablennamen in JavaScript in <dfn>camelCase</dfn>. Bei <dfn>camelCase</dfn> beginnt der erste Wortteil mit Kleinbuchstaben, und der erste Buchstabe jedes weiteren Wortteils wird großgeschrieben.

**Beispiele:**

```js
var someVariable;
var anotherVariableName;
var thisVariableNameIsSoLong;
```

# --instructions--

Ändern Sie die bestehenden Deklarationen und Zuweisungen so, dass ihre Namen <dfn>camelCase</dfn> verwenden.

Erstellen Sie keine neuen Variablen.

# --hints--

`studlyCapVar` sollte deklariert sein und den Wert `10` haben.

```js
assert(typeof studlyCapVar !== 'undefined' && studlyCapVar === 10);
```

`properCamelCase` sollte deklariert sein und den Wert des Strings `A String` haben.

```js
assert(
  typeof properCamelCase !== 'undefined' && properCamelCase === 'A String'
);
```

`titleCaseOver` sollte deklariert sein und den Wert `9000` haben.

```js
assert(typeof titleCaseOver !== 'undefined' && titleCaseOver === 9000);
```

`studlyCapVar` sollte in Deklaration und Zuweisung <dfn>camelCase</dfn> verwenden.

```js
assert(__helpers.removeJSComments(code).match(/studlyCapVar/g).length === 2);
```

`properCamelCase` sollte in Deklaration und Zuweisung <dfn>camelCase</dfn> verwenden.

```js
assert(__helpers.removeJSComments(code).match(/properCamelCase/g).length === 2);
```

`titleCaseOver` sollte in Deklaration und Zuweisung <dfn>camelCase</dfn> verwenden.

```js
assert(__helpers.removeJSComments(code).match(/titleCaseOver/g).length === 2);
```

# --seed--

## --seed-contents--

```js
// Variable declarations
var StUdLyCapVaR;
var properCamelCase;
var TitleCaseOver;

// Variable assignments
STUDLYCAPVAR = 10;
PRoperCAmelCAse = "A String";
tITLEcASEoVER = 9000;
```

# --solutions--

```js
var studlyCapVar;
var properCamelCase;
var titleCaseOver;

studlyCapVar = 10;
properCamelCase = "A String";
titleCaseOver = 9000;
```
