---
id: 56533eb9ac21ba0edf2244c8
title: Zugreifen auf Objekteigenschaften mit Klammernotation
challengeType: 1
forumTopicId: 16163
dashedName: accessing-object-properties-with-bracket-notation
---

# --description--

Die zweite Möglichkeit, auf die Eigenschaften eines Objekts zuzugreifen, ist die Klammernotation (`[]`). Wenn die Eigenschaft des Objekts, auf die Sie zugreifen möchten, ein Leerzeichen im Namen hat, müssen Sie die Klammernotation verwenden.

Sie können die Klammernotation jedoch auch bei Objekteigenschaften ohne Leerzeichen verwenden.

Hier ist ein Beispiel, wie Sie mit der Klammernotation eine Eigenschaft eines Objekts auslesen:

```js
const myObj = {
  "Space Name": "Kirk",
  "More Space": "Spock",
  "NoSpace": "USS Enterprise"
};

myObj["Space Name"];
myObj['More Space'];
myObj["NoSpace"];
```

`myObj["Space Name"]` wäre der String `Kirk`, `myObj['More Space']` wäre der String `Spock` und `myObj["NoSpace"]` wäre der String `USS Enterprise`.

Beachten Sie, dass Eigenschaftsnamen mit Leerzeichen in Anführungszeichen (einfach oder doppelt) stehen müssen.

# --instructions--

Lesen Sie die Werte der Eigenschaften `an entree` und `the drink` von `testObj` mit der Klammernotation aus und weisen Sie sie jeweils `entreeValue` und `drinkValue` zu.

# --hints--

`entreeValue` sollte ein String sein

```js
assert(typeof entreeValue === 'string');
```

Der Wert von `entreeValue` sollte der String `hamburger` sein

```js
assert(entreeValue === 'hamburger');
```

`drinkValue` sollte ein String sein

```js
assert(typeof drinkValue === 'string');
```

Der Wert von `drinkValue` sollte der String `water` sein

```js
assert(drinkValue === 'water');
```

Sie sollten die Klammernotation zweimal verwenden

```js
assert(__helpers.removeJSComments(code).match(/testObj\s*?\[('|")[^'"]+\1\]/g).length > 1);
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  "an entree": "hamburger",
  "my side": "veggies",
  "the drink": "water"
};

// Only change code below this line
const entreeValue = testObj;   // Change this line
const drinkValue = testObj;    // Change this line
```

# --solutions--

```js
const testObj = {
  "an entree": "hamburger",
  "my side": "veggies",
  "the drink": "water"
};
const entreeValue = testObj["an entree"];
const drinkValue = testObj['the drink'];
```
