---
id: 56533eb9ac21ba0edf2244bb
title: Wortlücken
challengeType: 1
forumTopicId: 18377
dashedName: word-blanks
---

# --description--

Sie erhalten Sätze mit einigen fehlenden Wörtern, wie Substantiven, Verben, Adjektiven und Adverbien. Sie füllen dann die fehlenden Teile mit Wörtern Ihrer Wahl so aus, dass der vollständige Satz Sinn ergibt.

Betrachten Sie diesen Satz:

```md
It was really ____, and we ____ ourselves ____.
```

Dieser Satz hat drei fehlende Teile – ein Adjektiv, ein Verb und ein Adverb – und wir können Wörter unserer Wahl hinzufügen, um ihn zu vervollständigen. Anschließend können wir den vollständigen Satz einer Variablen wie folgt zuweisen:

```js
const sentence = "It was really " + "hot" + ", and we " + "laughed" + " ourselves " + "silly" + ".";
```

# --instructions--

In dieser Aufgabe stellen wir Ihnen ein Substantiv, ein Verb, ein Adjektiv und ein Adverb zur Verfügung. Sie sollen einen vollständigen Satz bilden, indem Sie Wörter Ihrer Wahl zusammen mit den bereitgestellten Wörtern verwenden.

Sie müssen den String-Verkettungsoperator `+` verwenden, um einen neuen String zu erstellen, und dabei die bereitgestellten Variablen `myNoun`, `myAdjective`, `myVerb` und `myAdverb` verwenden. Anschließend weisen Sie den gebildeten String der Variablen `wordBlanks` zu. Die den Variablen zugewiesenen Wörter dürfen Sie nicht ändern.

Sie müssen auch auf Leerzeichen in Ihrem String achten, damit der endgültige Satz zwischen allen Wörtern Leerzeichen enthält. Das Ergebnis soll ein vollständiger Satz sein.

# --before-each--

```js
const removeAssignments = str => str
  .replace(/myNoun\s*=\s*["']dog["']/g, '')
  .replace(/myAdjective\s*=\s*["']big["']/g, '')
  .replace(/myVerb\s*=\s*["']ran["']/g, '')
  .replace(/myAdverb\s*=\s*["']quickly["']/g, '');
```

# --hints--

`wordBlanks` sollte ein String sein.

```js
assert(typeof wordBlanks === 'string');
```

Sie dürfen die Werte, die `myNoun`, `myVerb`, `myAdjective` oder `myAdverb` zugewiesen sind, nicht ändern.

```js
assert(
  myNoun === 'dog' &&
    myVerb === 'ran' &&
    myAdjective === 'big' &&
    myAdverb === 'quickly'
);
```

Sie sollten die Werte `dog`, `ran`, `big` oder `quickly` nicht direkt verwenden, um `wordBlanks` zu erstellen.

```js
const newCode = removeAssignments(__helpers.removeJSComments(code));
assert(
  !/dog/.test(newCode) &&
    !/ran/.test(newCode) &&
    !/big/.test(newCode) &&
    !/quickly/.test(newCode)
);
```

`wordBlanks` sollte alle Wörter enthalten, die den Variablen `myNoun`, `myVerb`, `myAdjective` und `myAdverb` zugewiesen sind, getrennt durch Nicht-Wort-Zeichen (und alle zusätzlichen Wörter Ihrer Wahl).

```js
assert(
  /\bdog\b/.test(wordBlanks) &&
    /\bbig\b/.test(wordBlanks) &&
    /\bran\b/.test(wordBlanks) &&
    /\bquickly\b/.test(wordBlanks)
);
```

# --seed--

## --seed-contents--

```js
const myNoun = "dog";
const myAdjective = "big";
const myVerb = "ran";
const myAdverb = "quickly";

// Only change code below this line
const wordBlanks = ""; // Change this line
// Only change code above this line
```

# --solutions--

```js
const myNoun = "dog";
const myAdjective = "big";
const myVerb = "ran";
const myAdverb = "quickly";

let wordBlanks = "Once there was a " + myNoun + " which was very " + myAdjective + ". ";
wordBlanks += "It " + myVerb + " " + myAdverb + " around the yard.";
```
