---
id: 56533eb9ac21ba0edf2244c9
title: Zugriff auf Objekteigenschaften mit Variablen
challengeType: 1
forumTopicId: 16165
dashedName: accessing-object-properties-with-variables
---

# --description--

Eine weitere Verwendung der Klammernotation bei Objekten ist der Zugriff auf eine Eigenschaft, die als Wert einer Variablen gespeichert ist. Dies kann sehr nützlich sein, um durch die Eigenschaften eines Objekts zu iterieren oder beim Zugriff auf eine Nachschlagetabelle.

Hier ist ein Beispiel, wie man eine Variable verwendet, um auf eine Eigenschaft zuzugreifen:

```js
const dogs = {
  Fido: "Mutt",
  Hunter: "Doberman",
  Snoopie: "Beagle"
};

const myDog = "Hunter";
const myBreed = dogs[myDog];
console.log(myBreed);
```

Der String `Doberman` würde in der Konsole angezeigt werden.

Beachten Sie, dass wir *keine* Anführungszeichen um den Variablennamen verwenden, wenn wir ihn zum Zugriff auf die Eigenschaft nutzen, weil wir den *Wert* der Variablen verwenden, nicht den *Namen*.

# --instructions--

Setzen Sie die `playerNumber`-Variable auf `16`. Verwenden Sie dann die Variable, um den Namen des Spielers nachzuschlagen und ihn `player` zuzuweisen.

# --hints--

`playerNumber` sollte eine Zahl sein

```js
assert(typeof playerNumber === 'number');
```

Die Variable `player` sollte ein String sein

```js
assert(typeof player === 'string');
```

Der Wert von `player` sollte der String `Montana` sein

```js
assert(player === 'Montana');
```

Sie sollten die Klammernotation verwenden, um auf `testObj` zuzugreifen

```js
assert(/testObj\s*?\[.*?\]/.test(__helpers.removeJSComments(code)));
```

Sie sollten den Wert `Montana` nicht direkt der Variablen `player` zuweisen.

```js
assert(!__helpers.removeJSComments(code).match(/player\s*=\s*"|\'\s*Montana\s*"|\'\s*;/gi));
```

Sie sollten die Variable `playerNumber` in Ihrer Klammernotation verwenden

```js
assert(/testObj\s*?\[\s*playerNumber\s*\]/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  12: "Namath",
  16: "Montana",
  19: "Unitas"
};

// Only change code below this line
const playerNumber = 42;  // Change this line
const player = testObj;   // Change this line
```

# --solutions--

```js
const testObj = {
  12: "Namath",
  16: "Montana",
  19: "Unitas"
};
const playerNumber = 16;
const player = testObj[playerNumber];
```
