---
id: 56533eb9ac21ba0edf2244cb
title: Komplexe Objekte manipulieren
challengeType: 1
forumTopicId: 18208
dashedName: manipulating-complex-objects
---

# --description--

Manchmal möchten Sie Daten in einer flexiblen <dfn>Datenstruktur</dfn> speichern. Ein JavaScript-Objekt ist eine Möglichkeit, flexible Daten zu handhaben. Sie erlauben beliebige Kombinationen von <dfn>Strings</dfn>, <dfn>Zahlen</dfn>, <dfn>Booleans</dfn>, <dfn>Arrays</dfn>, <dfn>Funktionen</dfn> und <dfn>Objekten</dfn>.

Hier ist ein Beispiel für eine komplexe Datenstruktur:

```js
const ourMusic = [
  {
    "artist": "Daft Punk",
    "title": "Homework",
    "release_year": 1997,
    "formats": [ 
      "CD", 
      "Cassette", 
      "LP"
    ],
    "gold": true
  }
];
```

Dies ist ein Array, das ein Objekt enthält. Das Objekt hat verschiedene Stücke von <dfn>Metadaten</dfn> über ein Album. Es enthält außerdem ein verschachteltes `formats` Array. Wenn Sie weitere Album-Datensätze hinzufügen möchten, können Sie dies tun, indem Sie Datensätze zum obersten Array hinzufügen. Objekte speichern Daten in einer Eigenschaft, die ein Schlüssel-Wert-Format hat. Im obigen Beispiel ist `"artist": "Daft Punk"` eine Eigenschaft, die den Schlüssel `artist` und den Wert `Daft Punk` hat.

**Hinweis:** Sie müssen nach jedem Objekt im Array ein Komma setzen, außer es ist das letzte Objekt im Array.

# --instructions--

Fügen Sie dem `myMusic` Array ein neues Album hinzu. Fügen Sie `artist` und `title` Strings, eine `release_year` Zahl und ein `formats` Array von Strings hinzu.

# --hints--

`myMusic` sollte ein Array sein

```js
assert(Array.isArray(myMusic));
```

`myMusic` sollte mindestens zwei Elemente haben

```js
assert(myMusic.length > 1);
```

Die Elemente im `myMusic` Array sollten Objekte sein

```js
myMusic.forEach(object => {assert.typeOf(object, 'object')})
```

Ihr Objekt in `myMusic` sollte mindestens 4 Eigenschaften haben

```js
myMusic.forEach(object => {assert(Object.keys(object).length > 3); });
```

Ihr Objekt in `myMusic` sollte die Eigenschaft `artist` enthalten, die ein String ist

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['artist']);
  assert.typeOf(object.artist, 'string')
})
```

Ihr Objekt in `myMusic` sollte die Eigenschaft `title` enthalten, die ein String ist

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['title']);
  assert.typeOf(object.title, 'string')
})
```

Ihr Objekt in `myMusic` sollte die Eigenschaft `release_year` enthalten, die eine Zahl ist

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['release_year']);
  assert.typeOf(object.release_year, 'number')
})
```

Ihr Objekt in `myMusic` sollte eine `formats` Eigenschaft enthalten, die ein Array ist

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['formats']);
  assert.typeOf(object.formats, 'array')
})
```

`formats` sollte ein Array von Strings mit mindestens zwei Elementen sein

```js
myMusic.forEach(object => {
  object.formats.forEach(format => {
    assert.typeOf(format, 'string')
  });
  assert.isAtLeast(object.formats.length, 2)
})
```

# --seed--

## --seed-contents--

```js
const myMusic = [
  {
    "artist": "Billy Joel",
    "title": "Piano Man",
    "release_year": 1973,
    "formats": [
      "CD",
      "8T",
      "LP"
    ],
    "gold": true
  }
];
```

# --solutions--

```js
const myMusic = [
  {
    "artist": "Billy Joel",
    "title": "Piano Man",
    "release_year": 1973,
    "formats": [
      "CS",
      "8T",
      "LP" ],
    "gold": true
  },
  {
    "artist": "ABBA",
    "title": "Ring Ring",
    "release_year": 1973,
    "formats": [
      "CS",
      "8T",
      "LP",
    "CD",
  ]
  }
];
```
