---
id: 56533eb9ac21ba0edf2244b7
title: Strings mit dem Plus-Operator verketten
challengeType: 1
forumTopicId: 16802
dashedName: concatenating-strings-with-plus-operator
---

# --description--

In JavaScript wird der `+`-Operator bei Verwendung mit einem `String`-Wert als <dfn>Verkettungsoperator</dfn> bezeichnet. Sie können einen neuen String aus anderen Strings erstellen, indem Sie diese <dfn>verketten</dfn>.

**Beispiel**

```js
'My name is Alan,' + ' I concatenate.'
```

**Hinweis:** Achten Sie auf Leerzeichen. Die Verkettung fügt keine Leerzeichen zwischen den verketteten Strings hinzu, daher müssen Sie diese selbst hinzufügen.

Beispiel:

```js
const ourStr = "I come first. " + "I come second.";
```

Der String `I come first. I come second.` würde in der Konsole angezeigt werden.
# --instructions--

Erstellen Sie `myStr` aus den Strings `This is the start.` und `This is the end.` mit dem `+`-Operator. Achten Sie darauf, ein Leerzeichen zwischen die beiden Strings zu setzen.

# --hints--

`myStr` sollte ein einzelnes Leerzeichen zwischen den beiden Strings enthalten.

```js
assert(/start\. This/.test(myStr));
```

`myStr` sollte den Wert des Strings `This is the start. This is the end.` haben.

```js
assert(myStr === 'This is the start. This is the end.');
```

Sie sollten den `+`-Operator verwenden, um `myStr` zu erstellen.

```js
assert(__helpers.removeJSComments(code).match(/(["']).*\1\s*\+\s*(["']).*\2/g));
```

`myStr` sollte mit dem Schlüsselwort `const` erstellt werden.

```js
assert(/const\s+myStr/.test(__helpers.removeJSComments(code)));
```

Sie sollten das Ergebnis der Variablen `myStr` zuweisen.

```js
assert(/myStr\s*=/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "This is the start. " + "This is the end.";
```
