---
id: 56533eb9ac21ba0edf2244a8
title: Werte mit dem Zuweisungsoperator speichern
challengeType: 1
forumTopicId: 18310
dashedName: storing-values-with-the-assignment-operator
---

# --description--

In JavaScript können Sie mit dem <dfn>Zuweisungsoperator</dfn> (`=`) einen Wert in einer Variablen speichern.

```js
myVariable = 5;
```

Dies weist den `Number`-Wert `5` `myVariable` zu.

Wenn rechts vom `=`-Operator Berechnungen stehen, werden diese ausgeführt, bevor der Wert der Variablen links vom Operator zugewiesen wird.

```js
var myVar;
myVar = 5;
```

Zuerst erstellt dieser Code eine Variable namens `myVar`. Dann weist der Code `5` `myVar` zu. Wenn `myVar` nun erneut im Code auftaucht, behandelt das Programm es so, als wäre es `5`.

# --instructions--

Weisen Sie der Variablen `7` den Wert `a` zu.

# --hints--

Sie sollten den Code oberhalb des angegebenen Kommentars nicht ändern.

```js
assert(/var a;/.test(__helpers.removeJSComments(code)));
```

`a` sollte den Wert 7 haben.

```js
assert(typeof a === 'number' && a === 7);
```

# --seed--

## --seed-contents--

```js
// Setup
var a;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
```
