---
id: 56533eb9ac21ba0edf2244d2
title: Vergleich mit dem Ungleichheitsoperator
challengeType: 1
forumTopicId: 16787
dashedName: comparison-with-the-inequality-operator
---

# --description--

Der Ungleichheitsoperator (`!=`) ist das Gegenteil des Gleichheitsoperators. Ungleich bedeutet nicht gleich. Der Ungleichheitsoperator gibt `false` zurück, wenn der Gleichheitsoperator `true` zurückgeben würde, und *umgekehrt*. Wie der Gleichheitsoperator wandelt auch der Ungleichheitsoperator die Datentypen der Werte beim Vergleichen um.

**Beispiele**

```js
1 !=  2    // true
1 != "1"   // false
1 != '1'   // false
1 != true  // false
0 != false // false
```

# --instructions--

Fügen Sie den Ungleichheitsoperator `!=` in die `if`-Anweisung ein, damit die Funktion den String `Not Equal` zurückgibt, wenn `val` nicht äquivalent zu `99` ist.

# --hints--

`testNotEqual(99)` sollte den String `Equal` zurückgeben

```js
assert(testNotEqual(99) === 'Equal');
```

`testNotEqual("99")` sollte den String `Equal` zurückgeben

```js
assert(testNotEqual('99') === 'Equal');
```

`testNotEqual(12)` sollte den String `Not Equal` zurückgeben

```js
assert(testNotEqual(12) === 'Not Equal');
```

`testNotEqual("12")` sollte den String `Not Equal` zurückgeben

```js
assert(testNotEqual('12') === 'Not Equal');
```

`testNotEqual("bob")` sollte den String `Not Equal` zurückgeben

```js
assert(testNotEqual('bob') === 'Not Equal');
```

Sie sollten den `!=`-Operator verwenden

```js
assert(__helpers.removeJSComments(code).match(/(?!!==)!=/));
```

# --seed--

## --seed-contents--

```js
// Setup
function testNotEqual(val) {
  if (val) { // Change this line
    return "Not Equal";
  }
  return "Equal";
}

testNotEqual(10);
```

# --solutions--

```js
function testNotEqual(val) {
  if (val != 99) {
    return "Not Equal";
  }
  return "Equal";
}
```
