---
id: 56533eb9ac21ba0edf2244af
title: Zusammengesetzte Zuweisung mit erweiterter Addition
challengeType: 1
forumTopicId: 16661
dashedName: compound-assignment-with-augmented-addition
---

# --description--

In der Programmierung ist es üblich, Zuweisungen zu verwenden, um den Inhalt einer Variablen zu ändern. Denken Sie daran, dass alles rechts vom Gleichheitszeichen zuerst ausgewertet wird, daher können wir sagen:

```js
myVar = myVar + 5;
```

um `5` zu `myVar` hinzuzufügen. Da dies ein so häufiges Muster ist, gibt es Operatoren, die sowohl eine mathematische Operation als auch eine Zuweisung in einem Schritt ausführen.

Einer dieser Operatoren ist der `+=`-Operator.

```js
let myVar = 1;
myVar += 5;
console.log(myVar);
```

`6` würde in der Konsole angezeigt werden.

# --instructions--

Konvertieren Sie die Zuweisungen für `a`, `b` und `c`, sodass sie den `+=`-Operator verwenden.

# --hints--

`a` sollte gleich `15` sein.

```js
assert(a === 15);
```

`b` sollte gleich `26` sein.

```js
assert(b === 26);
```

`c` sollte gleich `19` sein.

```js
assert(c === 19);
```

Sie sollten für jede Variable den `+=`-Operator verwenden.

```js
assert(__helpers.removeJSComments(code).match(/\+=/g).length === 3);
```

Sie dürfen den Code oberhalb des angegebenen Kommentars nicht ändern.

```js
assert(
  /let a = 3;/.test(__helpers.removeJSComments(code)) &&
    /let b = 17;/.test(__helpers.removeJSComments(code)) &&
    /let c = 12;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 3;
let b = 17;
let c = 12;

// Only change code below this line
a = a + 12;
b = 9 + b;
c = c + 7;
```

# --solutions--

```js
let a = 3;
let b = 17;
let c = 12;

a += 12;
b += 9;
c += 7;
```
