---
id: 56533eb9ac21ba0edf2244cd
title: Zugriff auf verschachtelte Arrays
challengeType: 1
forumTopicId: 16160
dashedName: accessing-nested-arrays
---

# --description--

Wie wir in früheren Beispielen gesehen haben, können Objekte sowohl verschachtelte Objekte als auch verschachtelte Arrays enthalten. Ähnlich wie beim Zugriff auf verschachtelte Objekte kann die Klammernotation für Arrays verkettet werden, um auf verschachtelte Arrays zuzugreifen.

Hier ist ein Beispiel, wie man auf ein verschachteltes Array zugreift:

```js
const ourPets = [
  {
    animalType: "cat",
    names: [
      "Meowzer",
      "Fluffy",
      "Kit-Cat"
    ]
  },
  {
    animalType: "dog",
    names: [
      "Spot",
      "Bowser",
      "Frankie"
    ]
  }
];

ourPets[0].names[1];
ourPets[1].names[0];
```

`ourPets[0].names[1]` wäre der String `Fluffy`, und `ourPets[1].names[0]` wäre der String `Spot`.

# --instructions--

Verwenden Sie Punkt- und Klammernotation, um die Variable `secondTree` auf das zweite Element im `list`-Array des zweiten Objekts im `myPlants`-Array zu setzen.

# --hints--

`secondTree` sollte dem String `pine` entsprechen.

```js
assert(secondTree === 'pine');
```

Ihr Code sollte Punkt- und Klammernotation verwenden, um auf `myPlants` zuzugreifen.

```js
assert(/=\s*myPlants\[1\].list\[1\]/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myPlants = [
  {
    type: "flowers",
    list: [
      "rose",
      "tulip",
      "dandelion"
    ]
  },
  {
    type: "trees",
    list: [
      "fir",
      "pine",
      "birch"
    ]
  }
];

const secondTree = "";
```

# --solutions--

```js
const myPlants = [
  {
    type: "flowers",
    list: [
      "rose",
      "tulip",
      "dandelion"
    ]
  },
  {
    type: "trees",
    list: [
      "fir",
      "pine",
      "birch"
    ]
  }
];

const secondTree = myPlants[1].list[1];
```
