---
id: 56533eb9ac21ba0edf2244de
title: Hinzufügen einer Standardoption in switch-Anweisungen
challengeType: 1
forumTopicId: 16653
dashedName: adding-a-default-option-in-switch-statements
---

# --description--

In einer `switch`-Anweisung können Sie möglicherweise nicht alle möglichen Werte als `case`-Anweisungen angeben. Stattdessen können Sie die `default`-Anweisung hinzufügen, die ausgeführt wird, wenn keine passenden `case`-Anweisungen gefunden werden. Denken Sie daran wie an die letzte `else`-Anweisung in einer `if/else`-Kette.

Eine `default`-Anweisung sollte der letzte Fall sein.

```js
switch (num) {
  case value1:
    statement1;
    break;
  case value2:
    statement2;
    break;
...
  default:
    defaultStatement;
    break;
}
```

# --instructions--

Schreiben Sie eine switch-Anweisung, um `answer` für die folgenden Bedingungen festzulegen:  
`a` - `apple`  
`b` - `bird`  
`c` - `cat`  
`default` - `stuff`

# --hints--

`switchOfStuff("a")` sollte den String `apple` zurückgeben

```js
assert(switchOfStuff('a') === 'apple');
```

`switchOfStuff("b")` sollte den String `bird` zurückgeben

```js
assert(switchOfStuff('b') === 'bird');
```

`switchOfStuff("c")` sollte den String `cat` zurückgeben

```js
assert(switchOfStuff('c') === 'cat');
```

`switchOfStuff("d")` sollte den String `stuff` zurückgeben

```js
assert(switchOfStuff('d') === 'stuff');
```

`switchOfStuff(4)` sollte den String `stuff` zurückgeben

```js
assert(switchOfStuff(4) === 'stuff');
```

Sie sollten keine `if`- oder `else`-Anweisungen verwenden

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

Sie sollten eine `default`-Anweisung verwenden

```js
assert(switchOfStuff('string-to-trigger-default-case') === 'stuff');
```

Sie sollten mindestens 3 `break`-Anweisungen haben

```js
assert(__helpers.removeJSComments(code).match(/break/g).length > 2);
```

# --seed--

## --seed-contents--

```js
function switchOfStuff(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

switchOfStuff(1);
```

# --solutions--

```js
function switchOfStuff(val) {
  let answer = "";

  switch(val) {
    case "a":
      answer = "apple";
      break;
    case "b":
      answer = "bird";
      break;
    case "c":
      answer = "cat";
      break;
    default:
      answer = "stuff";
  }
  return answer;
}
```
