---
id: 56533eb9ac21ba0edf2244bf
title: Lokaler Gültigkeitsbereich und Funktionen
challengeType: 1
forumTopicId: 18227
dashedName: local-scope-and-functions
---

# --description--

Variablen, die innerhalb einer Funktion deklariert werden, sowie die Parameter der Funktion haben einen <dfn>lokalen</dfn> Gültigkeitsbereich. Das bedeutet, sie sind nur innerhalb dieser Funktion sichtbar.

Hier ist eine Funktion `myTest` mit einer lokalen Variablen namens `loc`.

```js
function myTest() {
  const loc = "foo";
  console.log(loc);
}

myTest();
console.log(loc);
```

Der Funktionsaufruf `myTest()` wird den String `foo` in der Konsole anzeigen. Die Zeile `console.log(loc)` (außerhalb der Funktion `myTest`) wird einen Fehler werfen, da `loc` außerhalb der Funktion nicht definiert ist.

# --instructions--

Der Editor hat zwei `console.log`, die Ihnen helfen, zu sehen, was passiert. Überprüfen Sie die Konsole während des Codens, um zu sehen, wie sie sich verändert. Deklarieren Sie eine lokale Variable `myVar` innerhalb von `myLocalScope` und führen Sie die Tests aus.

**Hinweis:** Die Konsole zeigt weiterhin `ReferenceError: myVar is not defined` an, aber das führt nicht dazu, dass die Tests fehlschlagen.

# --hints--

Der Code sollte keine globale Variable `myVar` enthalten.

```js
function declared() {
  myVar;
}

assert.throws(declared, ReferenceError);
```

Sie sollten eine lokale Variable `myVar` hinzufügen.

```js
assert(
  /functionmyLocalScope\(\)\{.*(var|let|const)myVar[\s\S]*}/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

# --seed--

## --seed-contents--

```js
function myLocalScope() {
  // Only change code below this line

  console.log('inside myLocalScope', myVar);
}
myLocalScope();

// Run and check the console
// myVar is not defined outside of myLocalScope
console.log('outside myLocalScope', myVar);
```

# --solutions--

```js
function myLocalScope() {
  // Only change code below this line
  let myVar;
  console.log('inside myLocalScope', myVar);
}
myLocalScope();

// Run and check the console
// myVar is not defined outside of myLocalScope
console.log('outside myLocalScope', myVar);
```
