---
id: 56533eb9ac21ba0edf2244ed
title: Variablen an Strings anhängen
challengeType: 1
forumTopicId: 16656
dashedName: appending-variables-to-strings
---

# --description--

Genauso wie wir einen String über mehrere Zeilen aus String-<dfn>Literals</dfn> zusammensetzen können, können wir auch Variablen an einen String anhängen, und zwar mit dem Plus-Gleich-Operator (`+=`).

Beispiel:

```js
const anAdjective = "awesome!";
let ourStr = "freeCodeCamp is ";
ourStr += anAdjective;
```

`ourStr` hätte den Wert `freeCodeCamp is awesome!`.

# --instructions--

Setzen Sie `someAdjective` auf einen String mit mindestens 3 Zeichen und hängen Sie ihn mit dem `myStr`-Operator an `+=` an.

# --hints--

`someAdjective` sollte auf einen String mit mindestens 3 Zeichen gesetzt werden.

```js
assert(typeof someAdjective !== 'undefined' && someAdjective.length > 2);
```

Sie sollten `someAdjective` mit dem `myStr`-Operator an `+=` anhängen.

```js
assert(__helpers.removeJSComments(code).match(/myStr\s*\+=\s*someAdjective\s*/).length > 0);
```

# --seed--

## --seed-contents--

```js
// Change code below this line
const someAdjective = "";
let myStr = "Learning to code is ";
```

# --solutions--

```js
const someAdjective = "neat";
let myStr = "Learning to code is ";
myStr += someAdjective;
```
