---
id: 56533eb9ac21ba0edf2244c6
title: In die Schlange stellen
challengeType: 1
forumTopicId: 18307
dashedName: stand-in-line
---

# --description--

In der Informatik ist eine <dfn>queue</dfn> eine abstrakte <dfn>Datenstruktur</dfn>, in der Elemente in einer bestimmten Reihenfolge gehalten werden. Neue Elemente können am Ende der queue hinzugefügt werden, und alte Elemente werden von vorne aus der queue entfernt.

# --instructions--

Schreiben Sie eine Funktion `nextInLine`, die ein Array (`arr`) und eine Zahl (`item`) als Argumente entgegennimmt.

Fügen Sie die Zahl ans Ende des Arrays hinzu und entfernen Sie dann das erste Element des Arrays.

Die `nextInLine`-Funktion soll anschließend das entfernte Element zurückgeben.

# --hints--

`nextInLine([], 5)` sollte eine Zahl zurückgeben.

```js
assert.isNumber(nextInLine([], 5));
```

`nextInLine([], 1)` sollte `1` zurückgeben.

```js
assert(nextInLine([], 1) === 1);
```

`nextInLine([2], 1)` sollte `2` zurückgeben.

```js
assert(nextInLine([2], 1) === 2);
```

`nextInLine([5,6,7,8,9], 1)` sollte `5` zurückgeben.

```js
assert(nextInLine([5, 6, 7, 8, 9], 1) === 5);
```

# --seed--

## --seed-contents--

```js
function nextInLine(arr, item) {
  // Only change code below this line
  
  return item;
  // Only change code above this line
}

// Setup
let testArr = [1, 2, 3, 4, 5];

// Display code
console.log("Before: " + JSON.stringify(testArr));
console.log(nextInLine(testArr, 6));
console.log("After: " + JSON.stringify(testArr));
```

# --solutions--

```js
let testArr = [1, 2, 3, 4, 5];

function nextInLine(arr, item) {
    arr.push(item);
    return arr.shift();
}
```
