---
id: 56533eb9ac21ba0edf2244bc
title: Einkaufsliste
challengeType: 1
forumTopicId: 18280
dashedName: shopping-list
---

# --description--

Erstellen Sie eine Einkaufsliste in der Variablen `myList`. Die Liste soll ein mehrdimensionales Array sein, das mehrere Teilarrays enthält.

Das erste Element in jedem Teilarray soll einen String mit dem Namen des Artikels enthalten. Das zweite Element soll eine Zahl sein, die die Menge darstellt, z. B.

```js
["Chocolate Bar", 15]
```

Die Liste sollte mindestens 5 Teilarrays enthalten.

# --hints--

`myList` sollte ein Array sein.

```js
assert(Array.isArray(myList));
```

Die ersten Elemente in jedem Ihrer Teilarrays sollten alle Strings sein.

```js
assert(myList.every(elem => elem && typeof elem[0] === 'string'));
```

Die zweiten Elemente in jedem Ihrer Teilarrays sollten alle Zahlen sein.

```js
assert(myList.every(elem => elem && typeof elem[1] === 'number'));
```

Sie sollten mindestens 5 Artikel in Ihrer Liste haben.

```js
assert(myList.length > 4);
```

# --seed--

## --seed-contents--

```js
const myList = [];
```

# --solutions--

```js
const myList = [
  ["Candy", 10],
  ["Potatoes", 12],
  ["Eggs", 12],
  ["Catfood", 1],
  ["Toads", 9]
];
```
