---
id: 56533eb9ac21ba0edf2244b9
title: Strings mit Variablen erstellen
challengeType: 1
forumTopicId: 16805
dashedName: constructing-strings-with-variables
---

# --description--

Manchmal müssen Sie einen String erstellen. Mit dem Verkettungsoperator (`+`) können Sie eine oder mehrere Variablen in einen String einfügen, den Sie gerade aufbauen.

Beispiel:

```js
const ourName = "freeCodeCamp";
const ourStr = "Hello, our name is " + ourName + ", how are you?";
```

`ourStr` hätte den Wert des Strings `Hello, our name is freeCodeCamp, how are you?`.

# --instructions--

Setzen Sie `myName` auf einen String, der Ihrem Namen entspricht, und erstellen Sie `myStr` mit `myName` zwischen den Strings `My name is ` und ` and I am well!`.

# --hints--

`myName` sollte auf einen String mit mindestens 3 Zeichen gesetzt werden.

```js
assert(typeof myName !== 'undefined' && myName.length > 2);
```

Sie sollten zwei `+`-Operatoren verwenden, um `myStr` mit `myName` darin zu erstellen.

```js
assert(__helpers.removeJSComments(code).match(/["']\s*\+\s*myName\s*\+\s*["']/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const myName = "";
const myStr = "";
```

# --solutions--

```js
const myName = "Bob";
const myStr = "My name is " + myName + " and I am well!";
```
