---
id: 56533eb9ac21ba0edf2244ae
title: Den Rest in JavaScript finden
challengeType: 1
forumTopicId: 18184
dashedName: finding-a-remainder-in-javascript
---

# --description--

Der <dfn>Rest</dfn>-Operator `%` gibt den Rest der Division zweier Zahlen zurück.

**Beispiel**

<pre>
5 % 2 = 1
5 / 2 = 2 Rest 1
2 * 2 = 4
5 - 4 = 1
</pre>

**Verwendung**  
In der Mathematik kann überprüft werden, ob eine Zahl gerade oder ungerade ist, indem man den Rest der Division der Zahl durch `2` betrachtet. Gerade Zahlen haben einen Rest von `0`, während ungerade Zahlen einen Rest von `1` haben.

<pre>
17 % 2 = 1
48 % 2 = 0
</pre>

**Hinweis:** Der <dfn>Rest</dfn>-Operator wird manchmal fälschlicherweise als Modulus-Operator bezeichnet. Er ist dem Modulus sehr ähnlich, funktioniert aber bei negativen Zahlen nicht korrekt.

# --instructions--

Setzen Sie `remainder` auf den Rest von `11` geteilt durch `3` unter Verwendung des <dfn>Rest</dfn>-Operators (`%`).

# --hints--

Die Variable `remainder` sollte initialisiert werden

```js
assert(/(const|let|var)\s+?remainder/.test(__helpers.removeJSComments(code)));
```

Der Wert von `remainder` sollte `2` sein

```js
assert(remainder === 2);
```

Sie sollten den `%`-Operator verwenden

```js
assert(/\s+?remainder\s*?=\s*?.*%.*;?/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const remainder = 0;
```

# --solutions--

```js
const remainder = 11 % 3;
```
