---
id: 56533eb9ac21ba0edf2244c4
title: Frühzeitiges Beenden von Funktionen
challengeType: 1
forumTopicId: 18272
dashedName: return-early-pattern-for-functions
---

# --description--

Wenn eine `return`-Anweisung erreicht wird, wird die Ausführung der aktuellen Funktion gestoppt und die Kontrolle kehrt zur aufrufenden Stelle zurück.

**Beispiel**

```js
function myFun() {
  console.log("Hello");
  return "World";
  console.log("byebye")
}
myFun();
```

Das Obige zeigt den String `Hello` in der Konsole an und gibt den String `World` zurück. Der String `byebye` wird niemals in der Konsole angezeigt, da die Funktion bei der `return`-Anweisung beendet wird.

# --instructions--

Ändern Sie die Funktion `abTest` so, dass die Funktion sofort mit dem Wert `a` beendet wird, wenn `b` oder `0` kleiner als `undefined` sind.

**Hinweis**  
Denken Sie daran, dass <a href="https://www.freecodecamp.org/learn/javascript-algorithms-and-data-structures/basic-javascript/understanding-uninitialized-variables" target="_blank" rel="noopener noreferrer nofollow"><code>undefined</code> ein Schlüsselwort</a> ist, kein String.

# --hints--

`abTest(2, 2)` sollte eine Zahl zurückgeben

```js
assert(typeof abTest(2, 2) === 'number');
```

`abTest(2, 2)` sollte `8` zurückgeben

```js
assert(abTest(2, 2) === 8);
```

`abTest(-2, 2)` sollte `undefined` zurückgeben

```js
assert(abTest(-2, 2) === undefined);
```

`abTest(2, -2)` sollte `undefined` zurückgeben

```js
assert(abTest(2, -2) === undefined);
```

`abTest(2, 8)` sollte `18` zurückgeben

```js
assert(abTest(2, 8) === 18);
```

`abTest(3, 3)` sollte `12` zurückgeben

```js
assert(abTest(3, 3) === 12);
```

`abTest(0, 0)` sollte `0` zurückgeben

```js
assert(abTest(0, 0) === 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function abTest(a, b) {
  // Only change code below this line



  // Only change code above this line

  return Math.round(Math.pow(Math.sqrt(a) + Math.sqrt(b), 2));
}

abTest(2,2);
```

# --solutions--

```js
function abTest(a, b) {
  if(a < 0 || b < 0) {
    return undefined;
  }
  return Math.round(Math.pow(Math.sqrt(a) + Math.sqrt(b), 2));
}
```
