---
id: 56533eb9ac21ba0edf2244aa
title: Verständnis von nicht initialisierten Variablen
challengeType: 1
forumTopicId: 18335
dashedName: understanding-uninitialized-variables
---

# --description--

Wenn JavaScript-Variablen deklariert werden, haben sie einen Anfangswert von `undefined`. Wenn Sie eine mathematische Operation mit einer `undefined`-Variablen durchführen, ist Ihr Ergebnis `NaN`, was „Not a Number“ bedeutet. Wenn Sie einen String mit einer `undefined`-Variablen verketten, erhalten Sie einen String aus `undefined`.

# --instructions--

Initialisieren Sie die drei Variablen `a`, `b` und `c` mit `5`, `10` und `"I am a"`, damit sie nicht `undefined` sind.

# --hints--

`a` sollte definiert sein und einen Endwert von `6` haben.

```js
assert(typeof a === 'number' && a === 6);
```

`b` sollte definiert sein und einen Endwert von `15` haben.

```js
assert(typeof b === 'number' && b === 15);
```

`c` sollte kein `undefined` enthalten und einen Endwert des Strings `I am a String!` haben.

```js
assert(!/undefined/.test(c) && c === 'I am a String!');
```

Sie sollten den Code unterhalb des angegebenen Kommentars nicht ändern.

```js
assert(
  /a = a \+ 1;/.test(__helpers.removeJSComments(code)) &&
    /b = b \+ 5;/.test(__helpers.removeJSComments(code)) &&
    /c = c \+ " String!";/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
var a;
var b;
var c;
// Only change code above this line

a = a + 1;
b = b + 5;
c = c + " String!";
```

# --solutions--

```js
var a = 5;
var b = 10;
var c = "I am a";
a = a + 1;
b = b + 5;
c = c + " String!";
```
