---
id: 56533eb9ac21ba0edf2244b2
title: Zusammengesetzte Zuweisung mit erweiterter Division
challengeType: 1
forumTopicId: 16659
dashedName: compound-assignment-with-augmented-division
---

# --description--

Der `/=`-Operator dividiert eine Variable durch eine andere Zahl.

```js
myVar = myVar / 5;
```

Teilt `myVar` durch `5`. Dies kann umgeschrieben werden als:

```js
myVar /= 5;
```

# --instructions--

Konvertieren Sie die Zuweisungen für `a`, `b` und `c`, sodass der `/=`-Operator verwendet wird.

# --hints--

`a` sollte gleich `4` sein.

```js
assert(a === 4);
```

`b` sollte gleich `27` sein.

```js
assert(b === 27);
```

`c` sollte gleich `3` sein.

```js
assert(c === 3);
```

Sie sollten für jede Variable den `/=`-Operator verwenden.

```js
assert(__helpers.removeJSComments(code).match(/\/=/g).length === 3);
```

Sie dürfen den Code oberhalb des angegebenen Kommentars nicht verändern.

```js
assert(
  /let a = 48;/.test(__helpers.removeJSComments(code)) &&
    /let b = 108;/.test(__helpers.removeJSComments(code)) &&
    /let c = 33;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 48;
let b = 108;
let c = 33;

// Only change code below this line
a = a / 12;
b = b / 4;
c = c / 11;
```

# --solutions--

```js
let a = 48;
let b = 108;
let c = 33;

a /= 12;
b /= 4;
c /= 11;
```
