---
id: 56533eb9ac21ba0edf2244ad
title: Eine Zahl mit JavaScript dekrementieren
challengeType: 1
forumTopicId: 17558
dashedName: decrement-a-number-with-javascript
---

# --description--

Sie können eine Variable ganz einfach um eins <dfn>dekrementieren</dfn> oder verringern mit dem `--`-Operator.

```js
i--;
```

ist das Äquivalent zu

```js
i = i - 1;
```

**Hinweis:** Die gesamte Zeile wird zu `i--;`, wodurch das Gleichheitszeichen entfällt.

# --instructions--

Ändern Sie den Code so, dass der `--`-Operator auf `myVar` angewendet wird.

# --hints--

`myVar` sollte gleich `10` sein.

```js
assert(myVar === 10);
```

`myVar = myVar - 1;` sollte geändert werden.

```js
assert(!__helpers.removeJSComments(code).match(/myVar\s*=\s*myVar\s*[-]\s*1.*?;?/));
```

Sie sollten `myVar` nicht mit `10` zuweisen.

```js
assert(!__helpers.removeJSComments(code).match(/myVar\s*=\s*10.*?;?/));
```

Sie sollten den `--`-Operator auf `myVar` verwenden.

```js
assert(/[-]{2}\s*myVar|myVar\s*[-]{2}/.test(__helpers.removeJSComments(code)));
```

Sie sollten den Code oberhalb des angegebenen Kommentars nicht ändern.

```js
assert(/let myVar = 11;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let myVar = 11;

// Only change code below this line
myVar = myVar - 1;
```

# --solutions--

```js
let myVar = 11;
myVar--;
```
