---
id: 56533eb9ac21ba0edf2244cc
title: Zugriff auf verschachtelte Objekte
challengeType: 1
forumTopicId: 16161
dashedName: accessing-nested-objects
---

# --description--

Auf die Untereigenschaften von Objekten kann zugegriffen werden, indem man die Punktnotation oder Klammernotation miteinander verkettet.

Hier ist ein verschachteltes Objekt:

```js
const ourStorage = {
  "desk": {
    "drawer": "stapler"
  },
  "cabinet": {
    "top drawer": { 
      "folder1": "a file",
      "folder2": "secrets"
    },
    "bottom drawer": "soda"
  }
};

ourStorage.cabinet["top drawer"].folder2;
ourStorage.desk.drawer;
```

`ourStorage.cabinet["top drawer"].folder2` wäre der String `secrets`, und `ourStorage.desk.drawer` wäre der String `stapler`.

# --instructions--

Greifen Sie auf das `myStorage`-Objekt zu und weisen Sie den Inhalt der `glove box`-Eigenschaft der Variablen `gloveBoxContents` zu. Verwenden Sie für alle Eigenschaften, wo möglich, die Punktnotation, andernfalls die Klammernotation.

# --hints--

`gloveBoxContents` sollte dem String `maps` entsprechen.

```js
assert(gloveBoxContents === 'maps');
```

Ihr Code sollte die Punktnotation verwenden, wo immer es möglich ist, um auf `myStorage` zuzugreifen.

```js
assert.match(code, /myStorage\.car\.inside/);
```

`gloveBoxContents` sollte weiterhin mit `const` deklariert werden.

```js
assert.match(code, /const\s+gloveBoxContents\s*=/);
```

Sie sollten das `myStorage`-Objekt nicht verändern.

```js
const expectedMyStorage = {
  "car":{
    "inside":{
      "glove box":"maps",
      "passenger seat":"crumbs"
    },
    "outside":{
      "trunk":"jack"
    }
  }
};
assert.deepStrictEqual(myStorage, expectedMyStorage);
```

# --seed--

## --seed-contents--

```js
const myStorage = {
  "car": {
    "inside": {
      "glove box": "maps",
      "passenger seat": "crumbs"
     },
    "outside": {
      "trunk": "jack"
    }
  }
};

const gloveBoxContents = undefined;
```

# --solutions--

```js
const myStorage = {
  "car":{
    "inside":{
      "glove box":"maps",
      "passenger seat":"crumbs"
    },
    "outside":{
      "trunk":"jack"
    }
  }
};
const gloveBoxContents = myStorage.car.inside["glove box"];
```
