---
id: 56533eb9ac21ba0edf2244b8
title: Strings mit dem Plus-Gleich-Operator verketten
challengeType: 1
forumTopicId: 16803
dashedName: concatenating-strings-with-the-plus-equals-operator
---

# --description--

Wir können den `+=`-Operator auch verwenden, um einen String an das Ende einer bestehenden String-Variable zu <dfn>verketten</dfn>. Das ist sehr hilfreich, um einen langen String über mehrere Zeilen zu verteilen.

**Hinweis:** Achten Sie auf Leerzeichen. Die Verkettung fügt keine Leerzeichen zwischen den verketteten Strings hinzu, daher müssen Sie diese selbst einfügen.

Beispiel:

```js
let ourStr = "I come first. ";
ourStr += "I come second.";
```

`ourStr` hat jetzt den Wert des Strings `I come first. I come second.`.

# --instructions--

Bauen Sie `myStr` über mehrere Zeilen auf, indem Sie diese zwei Strings `This is the first sentence.` und `This is the second sentence.` mit dem `+=`-Operator verketten. Verwenden Sie den `+=`-Operator ähnlich wie im Beispiel gezeigt und fügen Sie unbedingt ein Leerzeichen zwischen die beiden Strings ein. Beginnen Sie damit, den ersten String `myStr` zuzuweisen, und fügen Sie dann den zweiten String hinzu.

# --hints--

`myStr` sollte ein einzelnes Leerzeichen zwischen den beiden Strings enthalten.

```js
assert(/sentence\. This/.test(myStr));
```

`myStr` sollte den Wert des Strings `This is the first sentence. This is the second sentence.` haben.

```js
assert(myStr === 'This is the first sentence. This is the second sentence.');
```

Sie sollten den `+=`-Operator verwenden, um `myStr` aufzubauen.

```js
assert(__helpers.removeJSComments(code).match(/myStr\s*\+=\s*(["']).*\1/g));
```

# --seed--

## --seed-contents--

```js
let myStr;
```

# --solutions--

```js
let myStr = "This is the first sentence. ";
myStr += "This is the second sentence.";
```
