---
id: 56533eb9ac21ba0edf2244cf
title: Plattensammlung
challengeType: 1
forumTopicId: 18261
dashedName: record-collection
---

# --description--

Sie erstellen eine Funktion, die bei der Wartung einer Musikalben-Sammlung hilft. Die Sammlung ist als Objekt organisiert, das mehrere Alben enthält, die ebenfalls Objekte sind. Jedes Album wird in der Sammlung mit einer eindeutigen `id` als Eigenschaftsname dargestellt. Innerhalb jedes Album-Objekts gibt es verschiedene Eigenschaften, die Informationen über das Album beschreiben. Nicht alle Alben haben vollständige Informationen.

Die `updateRecords`-Funktion nimmt 4 Argumente entgegen, die durch die folgenden Funktionsparameter dargestellt werden:

-   `records` – ein Objekt, das mehrere einzelne Alben enthält
-   `id` – eine Zahl, die ein bestimmtes Album im `records`-Objekt darstellt
-   `prop` – ein String, der den Namen der zu aktualisierenden Album-Eigenschaft darstellt
-   `value` – ein String, der die Informationen enthält, mit denen die Album-Eigenschaft aktualisiert wird

Vervollständigen Sie die Funktion unter Verwendung der folgenden Regeln, um das an die Funktion übergebene Objekt zu ändern.

-   Ihre Funktion muss immer das gesamte `records`-Objekt zurückgeben.
-   Wenn `value` ein leerer String ist, löschen Sie die angegebene `prop`-Eigenschaft vom Album.
-   Wenn `prop` nicht `tracks` ist und `value` kein leerer String ist, weisen Sie `value` der `prop` dieses Albums zu.
-   Wenn `prop` `tracks` ist und `value` kein leerer String ist, das Album aber keine `tracks`-Eigenschaft hat, erstellen Sie ein leeres Array und fügen `value` hinzu.
-   Wenn prop `tracks` ist und `value` kein leerer String ist, fügen Sie `value` ans Ende des vorhandenen `tracks`-Arrays des Albums hinzu.

**Hinweis:** Für die Tests wird eine Kopie des `recordCollection`-Objekts verwendet. Sie sollten das `recordCollection`-Objekt nicht direkt ändern.

# --before-each--

```js
const _recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};
```

# --hints--

Nach `updateRecords(recordCollection, 5439, "artist", "ABBA")` sollte `artist` der String `ABBA` sein

```js
assert(
  updateRecords(_recordCollection, 5439, 'artist', 'ABBA')[5439]['artist'] ===
    'ABBA'
);
```

Nach `updateRecords(recordCollection, 5439, "tracks", "Take a Chance on Me")` sollte `tracks` den String `Take a Chance on Me` als letztes und einziges Element haben.

```js
assert(
  updateRecords(_recordCollection, 5439, 'tracks', 'Take a Chance on Me') &&
  _recordCollection[5439]['tracks'].length === 1 &&
  _recordCollection[5439]['tracks'].pop() === 'Take a Chance on Me'
);
```

Nach `updateRecords(recordCollection, 2548, "artist", "")` sollte `artist` nicht gesetzt sein

```js
updateRecords(_recordCollection, 2548, 'artist', '');
assert(!_recordCollection[2548].hasOwnProperty('artist'));
```

Nach `updateRecords(recordCollection, 1245, "tracks", "Addicted to Love")` sollte `tracks` den String `Addicted to Love` als letztes Element haben.

```js
assert(
  updateRecords(_recordCollection, 1245, 'tracks', 'Addicted to Love')[1245][
    'tracks'
  ].pop() === 'Addicted to Love'
);
```

Nach `updateRecords(recordCollection, 2468, "tracks", "Free")` sollte `tracks` den String `1999` als erstes Element haben.

```js
assert(
  updateRecords(_recordCollection, 2468, 'tracks', 'Free')[2468][
    'tracks'
  ][0] === '1999'
);
```

Nach `updateRecords(recordCollection, 2548, "tracks", "")` sollte `tracks` nicht gesetzt sein

```js
updateRecords(_recordCollection, 2548, 'tracks', '');
assert(!_recordCollection[2548].hasOwnProperty('tracks'));
```

Nach `updateRecords(recordCollection, 1245, "albumTitle", "Riptide")` sollte `albumTitle` der String `Riptide` sein

```js
assert(
  updateRecords(_recordCollection, 1245, 'albumTitle', 'Riptide')[1245][
    'albumTitle'
  ] === 'Riptide'
);
```

# --seed--

## --seed-contents--

```js
// Setup
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};

// Only change code below this line
function updateRecords(records, id, prop, value) {
  return records;
}

updateRecords(recordCollection, 5439, 'artist', 'ABBA');
```

# --solutions--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};

// Only change code below this line
function updateRecords(records, id, prop, value) {
  if (value === '') delete records[id][prop];
  else if (prop === 'tracks') {
    records[id][prop] = records[id][prop] || [];
    records[id][prop].push(value);
  } else {
    records[id][prop] = value;
  }

  return records;
}
```
