---
id: 56533eb9ac21ba0edf2244ca
title: Objekte für Nachschlagen verwenden
challengeType: 1
forumTopicId: 18373
dashedName: using-objects-for-lookups
---

# --description--

Objekte können als Schlüssel/Wert-Speicher betrachtet werden, ähnlich einem Wörterbuch. Wenn Sie tabellarische Daten haben, können Sie ein Objekt verwenden, um Werte nachzuschlagen, anstatt eine `switch`-Anweisung oder eine `if/else`-Kette zu verwenden. Das ist besonders nützlich, wenn Sie wissen, dass Ihre Eingabedaten auf einen bestimmten Bereich beschränkt sind.

Hier ist ein Beispiel für ein Artikelobjekt:

```js
const article = {
  "title": "How to create objects in JavaScript",
  "link": "https://www.freecodecamp.org/news/a-complete-guide-to-creating-objects-in-javascript-b0e2450655e8/",
  "author": "Kaashan Hussain",
  "language": "JavaScript",
  "tags": "TECHNOLOGY",
  "createdAt": "NOVEMBER 28, 2018"
};

const articleAuthor = article["author"];
const articleLink = article["link"];

const value = "title";
const valueLookup = article[value];
```

`articleAuthor` ist der String `Kaashan Hussain`, `articleLink` ist der String `https://www.freecodecamp.org/news/a-complete-guide-to-creating-objects-in-javascript-b0e2450655e8/` und `valueLookup` ist der String `How to create objects in JavaScript`.

# --instructions--

Wandeln Sie die switch-Anweisung in ein Objekt namens `lookup` um. Verwenden Sie es, um `val` nachzuschlagen und den zugehörigen String der Variablen `result` zuzuweisen.

# --hints--

`phoneticLookup("alpha")` sollte den String `Adams` ergeben

```js
assert(phoneticLookup('alpha') === 'Adams');
```

`phoneticLookup("bravo")` sollte den String `Boston` ergeben

```js
assert(phoneticLookup('bravo') === 'Boston');
```

`phoneticLookup("charlie")` sollte den String `Chicago` ergeben

```js
assert(phoneticLookup('charlie') === 'Chicago');
```

`phoneticLookup("delta")` sollte den String `Denver` ergeben

```js
assert(phoneticLookup('delta') === 'Denver');
```

`phoneticLookup("echo")` sollte den String `Easy` ergeben

```js
assert(phoneticLookup('echo') === 'Easy');
```

`phoneticLookup("foxtrot")` sollte den String `Frank` ergeben

```js
assert(phoneticLookup('foxtrot') === 'Frank');
```

`phoneticLookup("")` sollte `undefined` entsprechen

```js
assert(typeof phoneticLookup('') === 'undefined');
```

Sie sollten die `return`-Anweisung nicht verändern

```js
assert(__helpers.removeJSComments(code).match(/return\sresult;/));
```

Sie sollten keine `case`-, `switch`- oder `if`-Anweisungen verwenden

```js
assert(
  !/case|switch|if/g.test(__helpers.removeJSComments(code).replace(/([/]{2}.*)|([/][*][^/*]*[*][/])/g, ''))
);
```

# --seed--

## --seed-contents--

```js
// Setup
function phoneticLookup(val) {
  let result = "";

  // Only change code below this line
  switch(val) {
    case "alpha":
      result = "Adams";
      break;
    case "bravo":
      result = "Boston";
      break;
    case "charlie":
      result = "Chicago";
      break;
    case "delta":
      result = "Denver";
      break;
    case "echo":
      result = "Easy";
      break;
    case "foxtrot":
      result = "Frank";
  }

  // Only change code above this line
  return result;
}

phoneticLookup("charlie");
```

# --solutions--

```js
function phoneticLookup(val) {
  let result = "";

  const lookup = {
    alpha: "Adams",
    bravo: "Boston",
    charlie: "Chicago",
    delta: "Denver",
    echo: "Easy",
    foxtrot: "Frank"
  };

  result = lookup[val];

  return result;
}
```
