---
id: 56533eb9ac21ba0edf2244dd
title: Aus vielen Optionen mit switch-Anweisungen auswählen
challengeType: 1
forumTopicId: 18277
dashedName: selecting-from-many-options-with-switch-statements
---

# --description--

Wenn Sie einen Wert mit vielen Optionen vergleichen müssen, können Sie eine <dfn>switch</dfn>-Anweisung verwenden. Eine `switch`-Anweisung vergleicht den Wert mit den <dfn>case</dfn>-Anweisungen, die verschiedene mögliche Werte definieren. Innerhalb eines <dfn>case</dfn>-Blocks können beliebige gültige JavaScript-Anweisungen ausgeführt werden, die vom ersten passenden `case`-Wert bis zum Auftreten eines `break` laufen.

Hier ist ein Beispiel für eine `switch`-Anweisung:

```js
switch (fruit) {
  case "apple":
    console.log("The fruit is an apple");
    break;
  case "orange":
    console.log("The fruit is an orange");
    break;
}
```

`case`-Werte werden mit strikter Gleichheit (`===`) getestet. Das `break` weist JavaScript an, die Ausführung der Anweisungen zu stoppen. Wenn das `break` weggelassen wird, wird die nächste Anweisung ausgeführt.

# --instructions--

Schreiben Sie eine switch-Anweisung, die `val` testet und `answer` für die folgenden Bedingungen setzt:  
`1` – `alpha`  
`2` – `beta`  
`3` – `gamma`  
`4` – `delta`

# --hints--

`caseInSwitch(1)` sollte den Wert des Strings `alpha` haben

```js
assert(caseInSwitch(1) === 'alpha');
```

`caseInSwitch(2)` sollte den Wert des Strings `beta` haben

```js
assert(caseInSwitch(2) === 'beta');
```

`caseInSwitch(3)` sollte den Wert des Strings `gamma` haben

```js
assert(caseInSwitch(3) === 'gamma');
```

`caseInSwitch(4)` sollte den Wert des Strings `delta` haben

```js
assert(caseInSwitch(4) === 'delta');
```

Sie sollten keine `if`- oder `else`-Anweisungen verwenden

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

Sie sollten mindestens 3 `break`-Anweisungen haben

```js
assert(__helpers.removeJSComments(code).match(/break/g).length > 2);
```

# --seed--

## --seed-contents--

```js
function caseInSwitch(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

caseInSwitch(1);
```

# --solutions--

```js
function caseInSwitch(val) {
  let answer = "";

  switch (val) {
    case 1:
      answer = "alpha";
      break;
    case 2:
      answer = "beta";
      break;
    case 3:
      answer = "gamma";
      break;
    case 4:
      answer = "delta";
  }
  return answer;
}
```
