---
id: 56533eb9ac21ba0edf2244be
title: Globaler Gültigkeitsbereich und Funktionen
challengeType: 1
forumTopicId: 18193
dashedName: global-scope-and-functions
---

# --description--

Im JavaScript bezeichnet der <dfn>Gültigkeitsbereich</dfn> die Sichtbarkeit von Variablen. Variablen, die außerhalb eines Funktionsblocks definiert sind, haben einen <dfn>globalen</dfn> Gültigkeitsbereich. Das bedeutet, sie sind überall in Ihrem JavaScript-Code sichtbar.

Variablen, die ohne die `let` oder `const` Schlüsselwörter deklariert werden, werden automatisch im `global` Gültigkeitsbereich erstellt. Dies kann an anderer Stelle in Ihrem Code oder beim erneuten Ausführen einer Funktion unbeabsichtigte Folgen haben. Sie sollten Ihre Variablen immer mit `let` oder `const` deklarieren.

# --instructions--

Deklarieren Sie mit `let` oder `const` eine globale Variable namens `myGlobal` außerhalb einer Funktion. Initialisieren Sie sie mit dem Wert `10`.

Weisen Sie innerhalb der Funktion `fun1` `5` `oopsGlobal` zu, ***ohne*** die Schlüsselwörter `var`, `let` oder `const` zu verwenden.

# --before-each--

```js
var oopsGlobal;
```

# --hints--

`myGlobal` sollte definiert sein

```js
assert(typeof myGlobal != 'undefined');
```

`myGlobal` sollte den Wert `10` haben

```js
assert(myGlobal === 10);
```

`myGlobal` sollte mit den Schlüsselwörtern `let` oder `const` deklariert werden

```js
assert(/(let|const)\s+myGlobal/.test(__helpers.removeJSComments(code)));
```

`oopsGlobal` sollte eine globale Variable sein und den Wert `5` haben

```js
fun1();
assert(typeof oopsGlobal != 'undefined');
```

# --seed--

## --seed-contents--

```js
// Declare the myGlobal variable below this line


function fun1() {
  // Assign 5 to oopsGlobal here

}

// Only change code above this line

function fun2() {
  let output = "";
  if (typeof myGlobal != "undefined") {
    output += "myGlobal: " + myGlobal;
  }
  if (typeof oopsGlobal != "undefined") {
    output += " oopsGlobal: " + oopsGlobal;
  }
  console.log(output);
}
```

# --solutions--

```js
const myGlobal = 10;

function fun1() {
  oopsGlobal = 5;
}

function fun2() {
  let output = "";
  if(typeof myGlobal != "undefined") {
    output += "myGlobal: " + myGlobal;
  }
  if(typeof oopsGlobal != "undefined") {
    output += " oopsGlobal: " + oopsGlobal;
  }
  console.log(output);
}
```
