---
id: 56533eb9ac21ba0edf2244c7
title: Zugriff auf Objekteigenschaften mit Punktnotation
challengeType: 1
forumTopicId: 16164
dashedName: accessing-object-properties-with-dot-notation
---

# --description--

Es gibt zwei Möglichkeiten, auf die Eigenschaften eines Objekts zuzugreifen: Punktnotation (`.`) und Klammernotation (`[]`), ähnlich wie bei einem Array.

Die Punktnotation verwenden Sie, wenn Sie den Namen der Eigenschaft, auf die Sie zugreifen möchten, im Voraus kennen.

Hier ist ein Beispiel für die Verwendung der Punktnotation (`.`), um eine Eigenschaft eines Objekts auszulesen:

```js
const myObj = {
  prop1: "val1",
  prop2: "val2"
};

const prop1val = myObj.prop1;
const prop2val = myObj.prop2;
```

`prop1val` hätte den Wert des Strings `val1`, und `prop2val` hätte den Wert des Strings `val2`.

# --instructions--

Lesen Sie die Eigenschaftswerte von `testObj` mit Punktnotation aus. Setzen Sie die Variable `hatValue` gleich der Eigenschaft `hat` des Objekts und die Variable `shirtValue` gleich der Eigenschaft `shirt` des Objekts.

# --hints--

`hatValue` sollte ein String sein

```js
assert(typeof hatValue === 'string');
```

Der Wert von `hatValue` sollte der String `ballcap` sein

```js
assert(hatValue === 'ballcap');
```

`shirtValue` sollte ein String sein

```js
assert(typeof shirtValue === 'string');
```

Der Wert von `shirtValue` sollte der String `jersey` sein

```js
assert(shirtValue === 'jersey');
```

Sie sollten die Punktnotation zweimal verwenden

```js
assert(__helpers.removeJSComments(code).match(/testObj\.\w+/g).length > 1);
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  hat: "ballcap",
  shirt: "jersey",
  shoes: "cleats"
};

// Only change code below this line
const hatValue = testObj;      // Change this line
const shirtValue = testObj;    // Change this line
```

# --solutions--

```js
const testObj = {
  hat: "ballcap",
  shirt: "jersey",
  shoes: "cleats"
};

const hatValue = testObj.hat;
const shirtValue = testObj.shirt;
```
