---
id: 56533eb9ac21ba0edf2244ac
title: Eine Zahl mit JavaScript erhöhen
challengeType: 1
forumTopicId: 18201
dashedName: increment-a-number-with-javascript
---

# --description--

Sie können eine Variable ganz einfach <dfn>erhöhen</dfn> oder um eins erhöhen mit dem `++`-Operator.

```js
i++;
```

entspricht

```js
i = i + 1;
```

**Hinweis:** Die gesamte Zeile wird zu `i++;`, wodurch das Gleichheitszeichen entfällt.

# --instructions--

Ändern Sie den Code so, dass der `++`-Operator auf `myVar` angewendet wird.

# --hints--

`myVar` sollte gleich `88` sein.

```js
assert(myVar === 88);
```

Sie sollten den Zuweisungsoperator nicht verwenden.

```js
assert(
  /let\s+myVar\s*=\s*87;\s*\/*.*\s*([+]{2}\s*myVar|myVar\s*[+]{2})/.test(__helpers.removeJSComments(code))
);
```

Sie sollten den `++`-Operator verwenden.

```js
assert(/[+]{2}\s*myVar|myVar\s*[+]{2}/.test(__helpers.removeJSComments(code)));
```

Sie sollten den Code oberhalb des angegebenen Kommentars nicht ändern.

```js
assert(/let myVar = 87;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let myVar = 87;

// Only change code below this line
myVar = myVar + 1;
```

# --solutions--

```js
let myVar = 87;
myVar++;
```
