---
id: 56533eb9ac21ba0edf2244b0
title: Zusammengesetzte Zuweisung mit erweiterter Subtraktion
challengeType: 1
forumTopicId: 16660
dashedName: compound-assignment-with-augmented-subtraction
---

# --description--

Wie der `+=`-Operator zieht `-=` eine Zahl von einer Variablen ab.

```js
myVar = myVar - 5;
```

zieht `5` von `myVar` ab. Das kann umgeschrieben werden als:

```js
myVar -= 5;
```

# --instructions--

Konvertieren Sie die Zuweisungen für `a`, `b` und `c`, sodass sie den `-=`-Operator verwenden.

# --hints--

`a` sollte gleich `5` sein.

```js
assert(a === 5);
```

`b` sollte gleich `-6` sein.

```js
assert(b === -6);
```

`c` sollte gleich `2` sein.

```js
assert(c === 2);
```

Sie sollten für jede Variable den `-=`-Operator verwenden.

```js
assert(__helpers.removeJSComments(code).match(/-=/g).length === 3);
```

Sie dürfen den Code oberhalb des angegebenen Kommentars nicht ändern.

```js
assert(
  /let a = 11;/.test(__helpers.removeJSComments(code)) && /let b = 9;/.test(__helpers.removeJSComments(code)) && /let c = 3;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 11;
let b = 9;
let c = 3;

// Only change code below this line
a = a - 6;
b = b - 15;
c = c - 1;
```

# --solutions--

```js
let a = 11;
let b = 9;
let c = 3;

a -= 6;
b -= 15;
c -= 1;
```
