---
id: 56533eb9ac21ba0edf2244d5
title: Vergleich mit dem größer-gleich-Operator
challengeType: 1
forumTopicId: 16785
dashedName: comparison-with-the-greater-than-or-equal-to-operator
---

# --description--

Der größer-gleich-Operator (`>=`) vergleicht die Werte von zwei Zahlen. Wenn die Zahl links größer oder gleich der Zahl rechts ist, gibt er `true` zurück. Andernfalls gibt er `false` zurück.

Wie der Gleichheitsoperator wandelt auch der größer-gleich-Operator beim Vergleichen Datentypen um.

**Beispiele**

```js
6   >=  6  // true
7   >= '3' // true
2   >=  3  // false
'7' >=  9  // false
```

# --instructions--

Fügen Sie den größer-gleich-Operator in die angegebenen Zeilen ein, damit die return-Anweisungen sinnvoll sind.

# --hints--

`testGreaterOrEqual(0)` sollte den String `Less than 10` zurückgeben

```js
assert(testGreaterOrEqual(0) === 'Less than 10');
```

`testGreaterOrEqual(9)` sollte den String `Less than 10` zurückgeben

```js
assert(testGreaterOrEqual(9) === 'Less than 10');
```

`testGreaterOrEqual(10)` sollte den String `10 or Over` zurückgeben

```js
assert(testGreaterOrEqual(10) === '10 or Over');
```

`testGreaterOrEqual(11)` sollte den String `10 or Over` zurückgeben

```js
assert(testGreaterOrEqual(11) === '10 or Over');
```

`testGreaterOrEqual(19)` sollte den String `10 or Over` zurückgeben

```js
assert(testGreaterOrEqual(19) === '10 or Over');
```

`testGreaterOrEqual(100)` sollte den String `20 or Over` zurückgeben

```js
assert(testGreaterOrEqual(100) === '20 or Over');
```

`testGreaterOrEqual(21)` sollte den String `20 or Over` zurückgeben

```js
assert(testGreaterOrEqual(21) === '20 or Over');
```

Sie sollten den `>=`-Operator mindestens zweimal verwenden

```js
assert(__helpers.removeJSComments(code).match(/val\s*>=\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testGreaterOrEqual(val) {
  if (val) {  // Change this line
    return "20 or Over";
  }

  if (val) {  // Change this line
    return "10 or Over";
  }

  return "Less than 10";
}

testGreaterOrEqual(10);
```

# --solutions--

```js
function testGreaterOrEqual(val) {
  if (val >= 20) {  // Change this line
    return "20 or Over";
  }

  if (val >= 10) {  // Change this line
    return "10 or Over";
  }

  return "Less than 10";
}
```
