---
id: 56533eb9ac21ba0edf2244ba
title: String-Unveränderlichkeit verstehen
challengeType: 1
forumTopicId: 18331
dashedName: understand-string-immutability
---

# --description--

In JavaScript sind `String`-Werte <dfn>unveränderlich</dfn>, was bedeutet, dass sie nach ihrer Erstellung nicht mehr verändert werden können.

Zum Beispiel erzeugt der folgende Code einen Fehler, weil der Buchstabe `B` im String `Bob` nicht in den Buchstaben `J` geändert werden kann:

```js
let myStr = "Bob";
myStr[0] = "J";
```

Beachten Sie, dass dies *nicht* bedeutet, dass `myStr` nicht neu zugewiesen werden könnte. Die einzige Möglichkeit, `myStr` zu ändern, besteht darin, ihm einen neuen Wert zuzuweisen, wie im folgenden Beispiel gezeigt:

```js
let myStr = "Bob";
myStr = "Job";
```

# --instructions--

Korrigieren Sie die Zuweisung an `myStr`, sodass sie den String-Wert von `Hello World` enthält, und verwenden Sie dabei die im obigen Beispiel gezeigte Vorgehensweise.

# --hints--

`myStr` sollte den Wert des Strings `Hello World` haben.

```js
assert(myStr === 'Hello World');
```

Sie sollten den Code oberhalb des angegebenen Kommentars nicht ändern.

```js
assert(/myStr = "Jello World"/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
let myStr = "Jello World";

// Only change code below this line
myStr[0] = "H"; // Change this line
// Only change code above this line
```

# --solutions--

```js
let myStr = "Jello World";
myStr = "Hello World";
```
