---
id: 56533eb9ac21ba0edf2244d8
title: Vergleiche mit dem logischen Und-Operator
challengeType: 1
forumTopicId: 16799
dashedName: comparisons-with-the-logical-and-operator
---

# --description--

Manchmal müssen Sie mehr als eine Bedingung gleichzeitig prüfen. Der <dfn>logische Und</dfn>-Operator (`&&`) gibt nur dann `true` zurück, wenn die <dfn>Operanden</dfn> links und rechts davon wahr sind.

Der gleiche Effekt kann erzielt werden, indem man eine `if`-Anweisung in eine andere `if`-Anweisung verschachtelt.

```js
if (num > 5) {
  if (num < 10) {
    return "Yes";
  }
}
return "No";
```

Dieser Code gibt `Yes` zurück, wenn `num` größer als `5` und kleiner als `10` ist. Dieselbe Logik kann mit dem <dfn>logischen Und</dfn>-Operator geschrieben werden.

```js
if (num > 5 && num < 10) {
  return "Yes";
}
return "No";
```

# --instructions--

Ersetzen Sie die zwei if-Anweisungen durch eine einzige Anweisung, die den `&&`-Operator verwendet und den String `Yes` zurückgibt, wenn `val` kleiner oder gleich `50` und größer oder gleich `25` ist. Andernfalls soll der String `No` zurückgegeben werden.

# --hints--

Sie sollten den `&&`-Operator einmal verwenden

```js
assert(__helpers.removeJSComments(code).match(/&&/g).length === 1);
```

Sie sollten nur eine `if`-Anweisung haben

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`testLogicalAnd(0)` sollte den String `No` zurückgeben

```js
assert(testLogicalAnd(0) === 'No');
```

`testLogicalAnd(24)` sollte den String `No` zurückgeben

```js
assert(testLogicalAnd(24) === 'No');
```

`testLogicalAnd(25)` sollte den String `Yes` zurückgeben

```js
assert(testLogicalAnd(25) === 'Yes');
```

`testLogicalAnd(30)` sollte den String `Yes` zurückgeben

```js
assert(testLogicalAnd(30) === 'Yes');
```

`testLogicalAnd(50)` sollte den String `Yes` zurückgeben

```js
assert(testLogicalAnd(50) === 'Yes');
```

`testLogicalAnd(51)` sollte den String `No` zurückgeben

```js
assert(testLogicalAnd(51) === 'No');
```

`testLogicalAnd(75)` sollte den String `No` zurückgeben

```js
assert(testLogicalAnd(75) === 'No');
```

`testLogicalAnd(80)` sollte den String `No` zurückgeben

```js
assert(testLogicalAnd(80) === 'No');
```

# --seed--

## --seed-contents--

```js
function testLogicalAnd(val) {
  // Only change code below this line

  if (val) {
    if (val) {
      return "Yes";
    }
  }

  // Only change code above this line
  return "No";
}

testLogicalAnd(10);
```

# --solutions--

```js
function testLogicalAnd(val) {
  if (val >= 25 && val <= 50) {
    return "Yes";
  }
  return "No";
}
```
