---
id: 56533eb9ac21ba0edf2244c3
title: Zuweisung mit einem Rückgabewert
challengeType: 1
forumTopicId: 16658
dashedName: assignment-with-a-returned-value
---

# --description--

Wenn Sie sich an unsere Diskussion über <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/storing-values-with-the-assignment-operator" target="_blank" rel="noopener noreferrer nofollow">Werte mit dem Zuweisungsoperator speichern</a> erinnern, wird alles rechts vom Gleichheitszeichen zuerst ausgewertet, bevor der Wert zugewiesen wird. Das bedeutet, wir können den Rückgabewert einer Funktion nehmen und ihn einer Variablen zuweisen.

Angenommen, wir haben eine Funktion `sum` definiert, die zwei Zahlen zusammenaddiert.

```js
ourSum = sum(5, 12);
```

Das Aufrufen der Funktion `sum` mit den Argumenten `5` und `12` ergibt einen Rückgabewert von `17`. Dieser Rückgabewert wird der Variablen `ourSum` zugewiesen.

# --instructions--

Rufen Sie die Funktion `processArg` mit dem Argument `7` auf und weisen Sie ihren Rückgabewert der Variablen `processed` zu.

# --hints--

`processed` sollte den Wert `2` haben

```js
assert(processed === 2);
```

Sie sollten `processArg` `processed` zuweisen

```js
assert(/processed\s*=\s*processArg\(\s*7\s*\)/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
let processed = 0;

function processArg(num) {
  return (num + 3) / 5;
}

// Only change code below this line

```

# --solutions--

```js
var processed = 0;

function processArg(num) {
  return (num + 3) / 5;
}

processed = processArg(7);
```
