---
id: 56592a60ddddeae28f7aa8e1
title: Mehrdimensionale Arrays mit Indexen zugreifen
challengeType: 1
forumTopicId: 16159
dashedName: access-multi-dimensional-arrays-with-indexes
---

# --description--

Eine Möglichkeit, sich ein <dfn>mehrdimensionales</dfn> Array vorzustellen, ist als ein *Array von Arrays*. Wenn Sie Klammern verwenden, um auf Ihr Array zuzugreifen, bezieht sich das erste Klammerpaar auf die Einträge im äußersten (ersten) Array, und jedes weitere Klammerpaar bezieht sich auf die nächste Ebene der Einträge darin.

**Beispiel**

```js
const arr = [
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9],
  [[10, 11, 12], 13, 14]
];

const subarray = arr[3];
const nestedSubarray = arr[3][0];
const element = arr[3][0][1];
```

In diesem Beispiel hat `subarray` den Wert `[[10, 11, 12], 13, 14]`,
`nestedSubarray` hat den Wert `[10, 11, 12]`, und `element` hat den Wert `11`.

**Hinweis:** Zwischen dem Array-Namen und den eckigen Klammern darf kein Leerzeichen stehen, wie bei `array [0][0]`, und auch dieses `array [0] [0]` ist nicht erlaubt. Obwohl JavaScript dies korrekt verarbeiten kann, könnte es andere Programmierende, die Ihren Code lesen, verwirren.

# --instructions--

Verwenden Sie die Klammernotation, um ein Element aus `myArray` auszuwählen, sodass `myData` gleich `8` ist.

# --hints--

`myData` sollte gleich `8` sein.

```js
assert(myData === 8);
```

Sie sollten die Klammernotation verwenden, um den richtigen Wert aus `myArray` auszulesen.

```js
assert(/myData=myArray\[2\]\[1\]/.test(__helpers.removeWhiteSpace(__helpers.removeJSComments(code))));
```

# --seed--

## --seed-contents--

```js
const myArray = [
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9],
  [[10, 11, 12], 13, 14],
];

const myData = myArray[0][0];
```

# --solutions--

```js
const myArray = [[1, 2, 3], [4, 5, 6], [7, 8, 9], [[10, 11, 12], 13, 14]];
const myData = myArray[2][1];
```
