---
id: 5675e877dbd60be8ad28edc6
title: Durchlaufen eines Arrays mit einer for-Schleife
challengeType: 1
forumTopicId: 18216
dashedName: iterate-through-an-array-with-a-for-loop
---

# --description--

Eine häufige Aufgabe in JavaScript ist es, den Inhalt eines Arrays zu durchlaufen. Eine Möglichkeit dafür ist eine `for`-Schleife. Dieser Code gibt jedes Element des Arrays `arr` in der Konsole aus:

```js
const arr = [10, 9, 8, 7, 6];

for (let i = 0; i < arr.length; i++) {
   console.log(arr[i]);
}
```

Denken Sie daran, dass Arrays nullbasiert indiziert sind, was bedeutet, dass der letzte Index des Arrays `length - 1` ist. Unsere Bedingung für diese Schleife ist `i < arr.length`, die die Schleife stoppt, wenn `i` gleich `length` ist. In diesem Fall ist die letzte Iteration `i === 4`, also wenn `i` gleich `arr.length - 1` wird und `6` in der Konsole ausgibt. Dann erhöht sich `i` auf `5`, und die Schleife wird beendet, weil `i < arr.length` `false` ist.

# --instructions--

Deklarieren und initialisieren Sie eine Variable `total` mit `0`. Verwenden Sie eine `for`-Schleife, um den Wert jedes Elements des `myArr`-Arrays zu `total` hinzuzufügen.

# --hints--

`total` sollte deklariert und mit 0 initialisiert werden.

```js
assert(__helpers.removeJSComments(code).match(/(var|let|const)\s*?total\s*=\s*0.*?;?/));
```

`total` sollte 20 sein.

```js
assert(total === 20);
```

Sie sollten eine `for`-Schleife verwenden, um `myArr` zu durchlaufen.

```js
assert(/for\s*\(/g.test(__helpers.removeJSComments(code)) && /myArr\s*\[/g.test(__helpers.removeJSComments(code)));
```

Sie sollten nicht versuchen, `total` direkt den Wert 20 zuzuweisen.

```js
assert(!__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/total[=+-]0*[1-9]+/gm));
```

# --seed--

## --seed-contents--

```js
// Setup
const myArr = [2, 3, 4, 5, 6];

// Only change code below this line

```

# --solutions--

```js
const myArr = [2, 3, 4, 5, 6];
let total = 0;

for (let i = 0; i < myArr.length; i++) {
  total += myArr[i];
}
```
