---
id: 5679ceb97cbaa8c51670a16b
title: Boolean-Werte aus Funktionen zurückgeben
challengeType: 1
forumTopicId: 18273
dashedName: returning-boolean-values-from-functions
---

# --description--

Sie erinnern sich vielleicht aus <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/comparison-with-the-equality-operator" target="_blank" rel="noopener noreferrer nofollow">Vergleich mit dem Gleichheitsoperator</a>, dass alle Vergleichsoperatoren einen Boolean `true` oder `false` Wert zurückgeben.

Manchmal verwenden Leute eine `if/else`-Anweisung, um einen Vergleich durchzuführen, so wie hier:

```js
function isEqual(a, b) {
  if (a === b) {
    return true;
  } else {
    return false;
  }
}
```

Aber es gibt eine bessere Methode. Da `===` `true` oder `false` zurückgibt, können wir das Ergebnis des Vergleichs direkt zurückgeben:

```js
function isEqual(a, b) {
  return a === b;
}
```

# --instructions--

Korrigieren Sie die Funktion `isLess`, indem Sie die `if/else`-Anweisungen entfernen.

# --hints--

`isLess(10, 15)` sollte `true` zurückgeben

```js
assert(isLess(10, 15) === true);
```

`isLess(15, 10)` sollte `false` zurückgeben

```js
assert(isLess(15, 10) === false);
```

Sie sollten keine `if`- oder `else`-Anweisungen verwenden

```js
assert(!/if|else/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function isLess(a, b) {
  // Only change code below this line
  if (a < b) {
    return true;
  } else {
    return false;
  }
  // Only change code above this line
}

isLess(10, 15);
```

# --solutions--

```js
function isLess(a, b) {
  return a < b;
}
```
