---
id: 56bbb991ad1ed5201cd392ce
title: Arrays mit der Methode unshift manipulieren
challengeType: 1
forumTopicId: 18239
dashedName: manipulate-arrays-with-unshift
---

# --description--

Sie können nicht nur Elemente mit `shift` vom Anfang eines Arrays entfernen, sondern auch Elemente mit `unshift` am Anfang eines Arrays hinzufügen, also Elemente vor das Array setzen.

`.unshift()` funktioniert genau wie `.push()`, aber anstatt das Element am Ende des Arrays hinzuzufügen, fügt `unshift()` das Element am Anfang des Arrays hinzu.

Beispiel:

```js
const ourArray = ["Stimpson", "J", "cat"];
ourArray.shift();
ourArray.unshift("Happy");
```

Nach dem `shift` hätte `ourArray` den Wert `["J", "cat"]`. Nach dem `unshift` hätte `ourArray` den Wert `["Happy", "J", "cat"]`.

# --instructions--

Fügen Sie `["Paul", 35]` mit `myArray` am Anfang der Variablen `unshift()` hinzu.

# --hints--

`myArray` sollte jetzt `[["Paul", 35], ["dog", 3]]` enthalten.

```js
assert(
  (function (d) {
    if (
      typeof d[0] === 'object' &&
      d[0][0] == 'Paul' &&
      d[0][1] === 35 &&
      d[1][0] != undefined &&
      d[1][0] == 'dog' &&
      d[1][1] != undefined &&
      d[1][1] == 3
    ) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["dog", 3]];
myArray.shift();

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["dog", 3]];
myArray.shift();
myArray.unshift(["Paul", 35]);
```
