---
id: 56bbb991ad1ed5201cd392ca
title: Zugriff auf Array-Daten mit Indexen
challengeType: 1
forumTopicId: 16158
dashedName: access-array-data-with-indexes
---

# --description--

Wir können auf die Daten innerhalb von Arrays mithilfe von <dfn>Indizes</dfn> zugreifen.

Array-Indizes werden in derselben Klammernotation geschrieben, die auch Strings verwenden, außer dass sie statt eines Zeichens einen Eintrag im Array angeben. Wie bei Strings verwenden Arrays eine <dfn>nullbasierte</dfn> Indexierung, sodass das erste Element in einem Array den Index `0` hat.

<br>

**Beispiel**

```js
const array = [50, 60, 70];
console.log(array[0]);
const data = array[1];
```

Das `console.log(array[0])` gibt `50` aus, und `data` hat den Wert `60`.

# --instructions--

Erstellen Sie eine Variable namens `myData` und setzen Sie sie auf den ersten Wert von `myArray` mithilfe der Klammernotation.

# --hints--

Die Variable `myData` sollte den ersten Wert von `myArray` enthalten.

```js
assert(
  (function () {
    if (
      typeof myArray !== 'undefined' &&
      typeof myData !== 'undefined' &&
      myArray[0] === myData
    ) {
      return true;
    } else {
      return false;
    }
  })()
);
```

Die Daten in der Variable `myArray` sollten mit Klammernotation aufgerufen werden.

```js
assert(
  (function () {
    if (__helpers.removeJSComments(code).match(/\s*=\s*myArray\[0\]/g)) {
      return true;
    } else {
      return false;
    }
  })()
);
```

# --seed--

## --seed-contents--

```js
const myArray = [50, 60, 70];


```

# --solutions--

```js
const myArray = [50, 60, 70];
const myData = myArray[0];
```
