---
id: 56bbb991ad1ed5201cd392cf
title: Wiederverwendbares JavaScript mit Funktionen schreiben
challengeType: 1
forumTopicId: 18378
dashedName: write-reusable-javascript-with-functions
---

# --description--

In JavaScript können wir unseren Code in wiederverwendbare Teile aufteilen, die <dfn>Funktionen</dfn> genannt werden.

Hier ist ein Beispiel für eine Funktion:

```js
function functionName() {
  console.log("Hello World");
}
```

Sie können diese Funktion aufrufen oder <dfn>ausführen</dfn>, indem Sie ihren Namen gefolgt von runden Klammern verwenden, so: `functionName();` Jedes Mal, wenn die Funktion aufgerufen wird, gibt sie die Nachricht `Hello World` in der Entwicklerkonsole aus. Der gesamte Code zwischen den geschweiften Klammern wird jedes Mal ausgeführt, wenn die Funktion aufgerufen wird.

# --instructions--

<ol>
  <li>
    Erstellen Sie eine Funktion namens <code>reusableFunction</code>, die den String <code>Hi World</code> in der Entwicklerkonsole ausgibt.
  </li>
  <li>
    Rufen Sie die Funktion auf.
  </li>
</ol>

# --before-each--

```js
function testConsole() {
  var originalConsole = console;
  var nativeLog = console.log;
  var hiWorldWasLogged = false;
  console.log = function (message) {
    if (message === 'Hi World') {
      console.warn(message);
      hiWorldWasLogged = true;
    }
    if (nativeLog.apply) {
      nativeLog.apply(originalConsole, arguments);
    } else {
      var nativeMsg = Array.prototype.slice.apply(arguments).join(' ');
      nativeLog(nativeMsg);
    }
  };
  reusableFunction();
  console.log = nativeLog;
  return hiWorldWasLogged;
}
```

# --hints--

`reusableFunction` sollte eine Funktion sein.

```js
assert(typeof reusableFunction === 'function');
```

Wenn `reusableFunction` aufgerufen wird, sollte der String `Hi World` in der Konsole ausgegeben werden.

```js
assert(testConsole());
```

Sie sollten `reusableFunction` aufrufen, sobald sie definiert ist.

```js
const functionStr = reusableFunction && __helpers.removeWhiteSpace(reusableFunction.toString());
const codeWithoutFunction = __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).replace(/reusableFunction\(\)\{/g, '');
assert(/reusableFunction\(\)/.test(codeWithoutFunction));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reusableFunction() {
  console.log("Hi World");
}
reusableFunction();
```
