---
id: 56bbb991ad1ed5201cd392d0
title: JavaScript-Objekte erstellen
challengeType: 1
forumTopicId: 16769
dashedName: build-javascript-objects
---

# --description--

Sie haben vielleicht schon einmal den Begriff `object` gehört.

Objekte sind ähnlich wie `arrays`, außer dass Sie anstelle von Indizes zum Zugreifen und Ändern ihrer Daten auf die Daten in Objekten über sogenannte `properties` zugreifen.

Objekte sind nützlich, um Daten strukturiert zu speichern, und können reale Objekte darstellen, wie zum Beispiel eine Katze.

Hier ist ein Beispiel für ein Katzenobjekt:

```js
const cat = {
  "name": "Whiskers",
  "legs": 4,
  "tails": 1,
  "enemies": ["Water", "Dogs"]
};
```

In diesem Beispiel sind alle Eigenschaften als Strings gespeichert, wie `name`, `legs` und `tails`. Sie können jedoch auch Zahlen als Eigenschaften verwenden. Sie können sogar die Anführungszeichen für einwortige String-Eigenschaften weglassen, wie folgt:

```js
const anotherObject = {
  make: "Ford",
  5: "five",
  "model": "focus"
};
```

Wenn Ihr Objekt jedoch nicht-string Eigenschaften enthält, wird JavaScript diese automatisch als Strings typumwandeln.

# --instructions--

Erstellen Sie ein Objekt, das einen Hund namens `myDog` darstellt und die Eigenschaften `name` (ein String), `legs`, `tails` und `friends` enthält.

Sie können diese Objekteigenschaften mit beliebigen Werten belegen, solange `name` ein String ist, `legs` und `tails` Zahlen sind und `friends` ein Array ist.

# --hints--

`myDog` sollte die Eigenschaft `name` enthalten, und diese sollte ein `string` sein.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('name') &&
      z.name !== undefined &&
      typeof z.name === 'string'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` sollte die Eigenschaft `legs` enthalten, und diese sollte ein `number` sein.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('legs') &&
      z.legs !== undefined &&
      typeof z.legs === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` sollte die Eigenschaft `tails` enthalten, und diese sollte ein `number` sein.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('tails') &&
      z.tails !== undefined &&
      typeof z.tails === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` sollte die Eigenschaft `friends` enthalten, und diese sollte ein `array` sein.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('friends') &&
      z.friends !== undefined &&
      Array.isArray(z.friends)
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` sollte nur alle angegebenen Eigenschaften enthalten.

```js
assert(
  (function (z) {
    return Object.keys(z).length === 4;
  })(myDog)
);
```

# --seed--

## --seed-contents--

```js
const myDog = {
  // Only change code below this line


  // Only change code above this line
};
```

# --solutions--

```js
const myDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```
