---
id: 56bbb991ad1ed5201cd392cd
title: Arrays mit der Methode shift manipulieren
challengeType: 1
forumTopicId: 18238
dashedName: manipulate-arrays-with-shift
---

# --description--

`pop()` entfernt immer das letzte Element eines Arrays. Was aber, wenn Sie das erste Element entfernen möchten?

Hier kommt `.shift()` ins Spiel. Es funktioniert genauso wie `.pop()`, entfernt jedoch das erste Element statt des letzten.

Beispiel:

```js
const ourArray = ["Stimpson", "J", ["cat"]];
const removedFromOurArray = ourArray.shift();
```

`removedFromOurArray` hätte den Wert des Strings `Stimpson`, und `ourArray` hätte `["J", ["cat"]]`.

# --instructions--

Verwenden Sie die `.shift()`-Funktion, um das erste Element von `myArray` zu entfernen, und weisen Sie den „herausgeschobenen“ Wert einer neuen Variablen `removedFromMyArray` zu.

# --hints--

`myArray` sollte jetzt `[["dog", 3]]` entsprechen.

```js
assert(
  (function (d) {
    if (d[0][0] == 'dog' && d[0][1] === 3 && d[1] == undefined) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

`removedFromMyArray` sollte `["John", 23]` enthalten.

```js
assert(
  (function (d) {
    if (
      d[0] == 'John' &&
      d[1] === 23 &&
      typeof removedFromMyArray === 'object'
    ) {
      return true;
    } else {
      return false;
    }
  })(removedFromMyArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["dog", 3]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["dog", 3]];

// Only change code below this line
const removedFromMyArray = myArray.shift();
```
