---
id: 56bbb991ad1ed5201cd392d2
title: Neue Eigenschaften zu einem JavaScript-Objekt hinzufügen
challengeType: 1
forumTopicId: 301169
dashedName: add-new-properties-to-a-javascript-object
---

# --description--

Sie können neuen Eigenschaften zu bestehenden JavaScript-Objekten auf dieselbe Weise hinzufügen, wie Sie sie ändern würden.

So würden wir eine `bark`-Eigenschaft zu `ourDog` hinzufügen:

```js
ourDog.bark = "bow-wow";
```

oder

```js
ourDog["bark"] = "bow-wow";
```

Wenn wir nun `ourDog.bark` auswerten, erhalten wir sein Bellen, `bow-wow`.

Beispiel:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};

ourDog.bark = "bow-wow";
```

# --instructions--

Fügen Sie eine `bark`-Eigenschaft zu `myDog` hinzu und setzen Sie sie auf einen Hundelaut, zum Beispiel „woof“. Sie können entweder die Punkt- oder die Klammernotation verwenden.

# --hints--

Sie sollten die Eigenschaft `bark` zu `myDog` hinzufügen.

```js
assert(myDog.bark !== undefined);
```

Sie sollten `bark` nicht zur Initialisierung von `myDog` hinzufügen.

```js
assert(!/bark[^\n]:/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};


```

# --solutions--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};
myDog.bark = "Woof Woof";
```
