---
id: 56bbb991ad1ed5201cd392cc
title: Arrays mit der Methode pop manipulieren
challengeType: 1
forumTopicId: 18236
dashedName: manipulate-arrays-with-pop
---

# --description--

Eine weitere Möglichkeit, die Daten in einem Array zu ändern, ist die Verwendung der `.pop()`-Funktion.

`.pop()` wird verwendet, um einen Wert vom Ende eines Arrays zu entfernen. Wir können diesen entfernten Wert speichern, indem wir ihn einer Variablen zuweisen. Anders gesagt, `.pop()` entfernt das letzte Element aus einem Array und gibt dieses Element zurück.

Jeder Eintragstyp kann von einem Array entfernt werden – Zahlen, Strings oder sogar verschachtelte Arrays.

```js
const threeArr = [1, 4, 6];
const oneDown = threeArr.pop();
console.log(oneDown);
console.log(threeArr);
```

Das erste `console.log` zeigt den Wert `6` an, und das zweite zeigt den Wert `[1, 4]` an.

# --instructions--

Verwenden Sie die `.pop()`-Funktion, um das letzte Element von `myArray` zu entfernen, und weisen Sie den entfernten Wert einer neuen Variablen `removedFromMyArray` zu.

# --hints--

`myArray` sollte nur `[["John", 23]]` enthalten.

```js
assert(
  (function (d) {
    if (d[0][0] == 'John' && d[0][1] === 23 && d[1] == undefined) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

Sie sollten `pop()` auf `myArray` anwenden.

```js
assert(/removedFromMyArray\s*=\s*myArray\s*.\s*pop\s*(\s*)/.test(__helpers.removeJSComments(code)));
```

`removedFromMyArray` sollte nur `["cat", 2]` enthalten.

```js
assert(
  (function (d) {
    if (d[0] == 'cat' && d[1] === 2 && d[2] == undefined) {
      return true;
    } else {
      return false;
    }
  })(removedFromMyArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["cat", 2]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["cat", 2]];
const removedFromMyArray = myArray.pop();
```
