---
id: 56bbb991ad1ed5201cd392d3
title: Eigenschaften aus einem JavaScript-Objekt löschen
challengeType: 1
forumTopicId: 17560
dashedName: delete-properties-from-a-javascript-object
---

# --description--

Wir können Eigenschaften auch so aus Objekten löschen:

```js
delete ourDog.bark;
```

Beispiel:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"],
  "bark": "bow-wow"
};

delete ourDog.bark;
```

Nach der letzten oben gezeigten Zeile sieht `ourDog` so aus:

```js
{
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
}
```

# --instructions--

Löschen Sie die Eigenschaft `tails` aus `myDog`. Sie können entweder die Punkt- oder die Klammernotation verwenden.

# --hints--

Sie sollten die Eigenschaft `tails` aus `myDog` löschen.

```js
assert(typeof myDog === 'object' && myDog.tails === undefined);
```

Sie sollten die `myDog` Einrichtung nicht verändern.

```js
assert(__helpers.removeJSComments(code).match(/"tails": 1/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"],
  "bark": "woof"
};

// Only change code below this line

```

# --solutions--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"],
  "bark": "woof"
};
delete myDog.tails;
```
