---
id: 56bbb991ad1ed5201cd392d1
title: Aktualisieren von Objekteigenschaften
challengeType: 1
forumTopicId: 18336
dashedName: updating-object-properties
---

# --description--

Nachdem Sie ein JavaScript-Objekt erstellt haben, können Sie seine Eigenschaften jederzeit aktualisieren, genau wie jede andere Variable. Sie können entweder die Punkt- oder die Klammernotation verwenden, um zu aktualisieren.

Zum Beispiel schauen wir uns `ourDog` an:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```

Da er ein besonders glücklicher Hund ist, ändern wir seinen Namen in den String `Happy Camper`. So aktualisieren wir die Eigenschaft name seines Objekts: `ourDog.name = "Happy Camper";` oder `ourDog["name"] = "Happy Camper";` Wenn wir jetzt `ourDog.name` auswerten, erhalten wir statt `Camper` seinen neuen Namen, `Happy Camper`.

# --instructions--

Aktualisieren Sie die Eigenschaft name des Objekts `myDog`. Ändern Sie ihren Namen von `Coder` zu `Happy Coder`. Sie können entweder die Punkt- oder die Klammernotation verwenden.

# --hints--

Sie sollten die Eigenschaft `myDog` von `name` auf den String `Happy Coder` setzen.

```js
assert(/happy coder/gi.test(myDog.name));
```

Sie sollten die Definition von `myDog` nicht bearbeiten.

```js
assert(/"name": "Coder"/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const myDog = {
  "name": "Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};

// Only change code below this line

```

# --solutions--

```js
const myDog = {
  "name": "Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};
myDog.name = "Happy Coder";
```
