---
id: 587d7b7e367417b2b2512b21
title: Mehrere bedingte (ternäre) Operatoren verwenden
challengeType: 1
forumTopicId: 301179
dashedName: use-multiple-conditional-ternary-operators
---

# --description--

In der vorherigen Aufgabe haben Sie einen einzelnen bedingten Operator verwendet. Sie können sie auch verketten, um mehrere Bedingungen zu prüfen.

Die folgende Funktion verwendet `if`, `else if` und `else`-Anweisungen, um mehrere Bedingungen zu prüfen:

```js
function findGreaterOrEqual(a, b) {
  if (a === b) {
    return "a and b are equal";
  }
  else if (a > b) {
    return "a is greater";
  }
  else {
    return "b is greater";
  }
}
```

Die obige Funktion kann mit mehreren bedingten Operatoren neu geschrieben werden:

```js
function findGreaterOrEqual(a, b) {
  return (a === b) ? "a and b are equal" 
    : (a > b) ? "a is greater" 
    : "b is greater";
}
```

Es gilt als bewährte Vorgehensweise, mehrere bedingte Operatoren so zu formatieren, dass jede Bedingung in einer eigenen Zeile steht, wie oben gezeigt. Die Verwendung mehrerer bedingter Operatoren ohne richtige Einrückung kann Ihren Code schwer lesbar machen. Zum Beispiel:

```js
function findGreaterOrEqual(a, b) {
  return (a === b) ? "a and b are equal" : (a > b) ? "a is greater" : "b is greater";
}
```

# --instructions--

Verwenden Sie in der `checkSign`-Funktion mehrere bedingte Operatoren – entsprechend dem empfohlenen Format in `findGreaterOrEqual` – um zu prüfen, ob eine Zahl positiv, negativ oder null ist. Die Funktion soll `positive`, `negative` oder `zero` zurückgeben.

# --hints--

`checkSign` sollte mehrere bedingte Operatoren verwenden

```js
assert(/.+?\s*?\?\s*?.+?\s*?:\s*?.+?\s*?\?\s*?.+?\s*?:\s*?.+?/gi.test(__helpers.removeJSComments(code)));
```

`checkSign(10)` sollte den String `positive` zurückgeben. Beachten Sie, dass die Groß- und Kleinschreibung wichtig ist

```js
assert(checkSign(10) === 'positive');
```

`checkSign(-12)` sollte den String `negative` zurückgeben. Beachten Sie, dass die Groß- und Kleinschreibung wichtig ist

```js
assert(checkSign(-12) === 'negative');
```

`checkSign(0)` sollte den String `zero` zurückgeben. Beachten Sie, dass die Groß- und Kleinschreibung wichtig ist

```js
assert(checkSign(0) === 'zero');
```

# --seed--

## --seed-contents--

```js
function checkSign(num) {

}

checkSign(10);
```

# --solutions--

```js
function checkSign(num) {
  return (num > 0) ? 'positive' : (num < 0) ? 'negative' : 'zero';
}
```
