---
id: 587d7b7e367417b2b2512b22
title: Die parseInt-Funktion mit einer Basis verwenden
challengeType: 1
forumTopicId: 301182
dashedName: use-the-parseint-function-with-a-radix
---

# --description--

Die `parseInt()`-Funktion parst einen String und gibt eine ganze Zahl zurück. Sie nimmt ein zweites Argument für die Basis, das die Zahlensystembasis des Strings angibt. Die Basis kann eine ganze Zahl zwischen 2 und 36 sein.

Der Funktionsaufruf sieht so aus:

```js
parseInt(string, radix);
```

Und hier ist ein Beispiel:

```js
const a = parseInt("11", 2);
```

Die Variable radix gibt an, dass `11` im Binärsystem, also Basis 2, vorliegt. Dieses Beispiel wandelt den String `11` in eine ganze Zahl `3` um.

# --instructions--

Verwenden Sie `parseInt()` in der `convertToInteger`-Funktion, damit sie eine Binärzahl in eine ganze Zahl umwandelt und zurückgibt.

# --hints--

`convertToInteger` sollte die `parseInt()`-Funktion verwenden

```js
assert(/parseInt/g.test(__helpers.removeJSComments(code)));
```

`convertToInteger("10011")` sollte eine Zahl zurückgeben

```js
assert(typeof convertToInteger('10011') === 'number');
```

`convertToInteger("10011")` sollte 19 zurückgeben

```js
assert(convertToInteger('10011') === 19);
```

`convertToInteger("111001")` sollte 57 zurückgeben

```js
assert(convertToInteger('111001') === 57);
```

`convertToInteger("JamesBond")` sollte `NaN` zurückgeben

```js
assert.isNaN(convertToInteger('JamesBond'));
```

# --seed--

## --seed-contents--

```js
function convertToInteger(str) {

}

convertToInteger("10011");
```

# --solutions--

```js
function convertToInteger(str) {
  return parseInt(str, 2);
}
```
