---
id: 587d7b7e367417b2b2512b23
title: Die parseInt-Funktion verwenden
challengeType: 1
forumTopicId: 301183
dashedName: use-the-parseint-function
---

# --description--

Die `parseInt()`-Funktion parst einen String und gibt eine ganze Zahl zurück. Hier ist ein Beispiel:

```js
const a = parseInt("007");
```

Die obige Funktion wandelt den String `007` in die ganze Zahl `7` um. Wenn das erste Zeichen im String nicht in eine Zahl umgewandelt werden kann, gibt sie `NaN` zurück.

# --instructions--

Verwenden Sie `parseInt()` in der `convertToInteger`-Funktion, damit sie den Eingabe-String `str` in eine ganze Zahl umwandelt und zurückgibt.

# --hints--

`convertToInteger` sollte die `parseInt()`-Funktion verwenden

```js
assert(/parseInt/g.test(__helpers.removeJSComments(code)));
```

`convertToInteger("56")` sollte eine Zahl zurückgeben

```js
assert(typeof convertToInteger('56') === 'number');
```

`convertToInteger("56")` sollte 56 zurückgeben

```js
assert(convertToInteger('56') === 56);
```

`convertToInteger("77")` sollte 77 zurückgeben

```js
assert(convertToInteger('77') === 77);
```

`convertToInteger("JamesBond")` sollte `NaN` zurückgeben

```js
assert.isNaN(convertToInteger('JamesBond'));
```

# --seed--

## --seed-contents--

```js
function convertToInteger(str) {

}

convertToInteger("56");
```

# --solutions--

```js
function convertToInteger(str) {
  return parseInt(str);
}
```
