---
id: 598e8944f009e646fc236146
title: "Verstehen des undefined-Werts, der von einer Funktion zurückgegeben wird"
challengeType: 1
forumTopicId: 301177
dashedName: understanding-undefined-value-returned-from-a-function
---

# --description--

Eine Funktion kann die `return`-Anweisung enthalten, muss sie aber nicht. Falls die Funktion keine `return`-Anweisung hat, verarbeitet die Funktion beim Aufrufen den inneren Code, aber der zurückgegebene Wert ist `undefined`.

**Beispiel**

```js
let sum = 0;

function addSum(num) {
  sum = sum + num;
}

addSum(3);
```

`addSum` ist eine Funktion ohne `return`-Anweisung. Die Funktion ändert die globale Variable `sum`, aber der zurückgegebene Wert der Funktion ist `undefined`.

# --instructions--

Erstellen Sie eine Funktion `addFive` ohne Argumente. Diese Funktion addiert 5 zur Variable `sum`, aber ihr zurückgegebener Wert ist `undefined`.

# --hints--

`addFive` sollte eine Funktion sein.

```js
assert(typeof addFive === 'function');
```

Nachdem beide Funktionen ausgeführt wurden, sollte `sum` gleich `8` sein.

```js
assert(sum === 8);
```

Der zurückgegebene Wert von `addFive` sollte `undefined` sein.

```js
assert(addFive() === undefined);
```

Innerhalb der Funktion `addFive` sollten Sie `5` zur Variable `sum` hinzufügen.

```js
assert(
  __helpers.removeWhiteSpace(addFive.toString()).match(/sum=sum\+5|sum\+=5/)
);
```

# --seed--

## --seed-contents--

```js
// Setup
let sum = 0;

function addThree() {
  sum = sum + 3;
}

// Only change code below this line


// Only change code above this line

addThree();
addFive();
```

# --solutions--

```js
let sum = 0;

function addThree() {
  sum = sum + 3;
}

function addFive() {
  sum = sum + 5;
}

addThree();
addFive();
```
