---
id: 5cc0bd7a49b71cb96132e54c
title: "Verwenden Sie Rekursion, um einen Zahlenbereich zu erstellen"
challengeType: 1
forumTopicId: 301180
dashedName: use-recursion-to-create-a-range-of-numbers
---

# --description--

Aufbauend auf der vorherigen Aufgabe erhalten Sie hier eine weitere Gelegenheit, eine rekursive Funktion zur Lösung eines Problems zu erstellen.

# --instructions--

Wir haben eine Funktion namens `rangeOfNumbers` mit zwei Parametern definiert. Die Funktion soll ein Array ganzer Zahlen zurückgeben, das mit einer Zahl beginnt, die durch den Parameter `startNum` dargestellt wird, und mit einer Zahl endet, die durch den Parameter `endNum` dargestellt wird. Die Startzahl ist immer kleiner oder gleich der Endzahl. Ihre Funktion muss Rekursion verwenden, indem sie sich selbst aufruft, und darf keine Schleifen verwenden. Sie soll auch für Fälle funktionieren, in denen sowohl `startNum` als auch `endNum` gleich sind.

# --hints--

Ihre Funktion sollte ein Array zurückgeben.

```js
assert(Array.isArray(rangeOfNumbers(5, 10)));
```

Ihr Code darf keine Schleifensyntax verwenden (`for` oder `while` oder höherwertige Funktionen wie `forEach`, `map`, `filter` oder `reduce`).

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

`rangeOfNumbers` sollte Rekursion verwenden (sich selbst aufrufen), um diese Aufgabe zu lösen.

```js
assert(
  rangeOfNumbers.toString().match(/rangeOfNumbers\s*\(.+\)/)
);
```

`rangeOfNumbers(1, 5)` sollte `[1, 2, 3, 4, 5]` zurückgeben.

```js
assert.deepStrictEqual(rangeOfNumbers(1, 5), [1, 2, 3, 4, 5]);
```

`rangeOfNumbers(6, 9)` sollte `[6, 7, 8, 9]` zurückgeben.

```js
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

`rangeOfNumbers(4, 4)` sollte `[4]` zurückgeben.

```js
assert.deepStrictEqual(rangeOfNumbers(4, 4), [4]);
```

Globale Variablen sollten nicht verwendet werden, um das Array zwischenzuspeichern.

```js
rangeOfNumbers(1, 3)
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

# --seed--

## --seed-contents--

```js
function rangeOfNumbers(startNum, endNum) {
  return [];
};
```

# --solutions--

```js
function rangeOfNumbers(startNum, endNum) {
  if (endNum - startNum === 0) {
    return [startNum];
  } else {
    const numbers = rangeOfNumbers(startNum, endNum - 1);
    numbers.push(endNum);
    return numbers;
  }
}
```
