---
id: 5cfa3679138e7d9595b9d9d4
title: Schleifen durch Rekursion ersetzen
challengeType: 1
videoUrl: >-
  https://www.freecodecamp.org/news/how-recursion-works-explained-with-flowcharts-and-a-video-de61f40cb7f9/
forumTopicId: 301175
dashedName: replace-loops-using-recursion
---

# --description--

Rekursion ist das Konzept, dass eine Funktion in Bezug auf sich selbst ausgedrückt werden kann. Um dies besser zu verstehen, denken Sie zunächst an die folgende Aufgabe: multiplizieren Sie die ersten `n` Elemente eines Arrays, um das Produkt dieser Elemente zu erhalten. Mit einer `for` Schleife könnten Sie das so machen:

```js
  function multiply(arr, n) {
    let product = 1;
    for (let i = 0; i < n; i++) {
      product *= arr[i];
    }
    return product;
  }
```

Beachten Sie jedoch, dass `multiply(arr, n) == multiply(arr, n - 1) * arr[n - 1]`. Das bedeutet, Sie können `multiply` in Bezug auf sich selbst umschreiben und brauchen niemals eine Schleife zu verwenden.

```js
  function multiply(arr, n) {
    if (n <= 0) {
      return 1;
    } else {
      return multiply(arr, n - 1) * arr[n - 1];
    }
  }
```

Die rekursive Version von `multiply` zerlegt sich folgendermaßen. Im <dfn>Basisfall</dfn>, wenn `n <= 0`, gibt sie 1 zurück. Für größere Werte von `n` ruft sie sich selbst auf, aber mit `n - 1`. Dieser Funktionsaufruf wird auf die gleiche Weise ausgewertet, ruft `multiply` erneut auf, bis `n <= 0` erreicht ist. An diesem Punkt können alle Funktionen zurückkehren und die ursprüngliche `multiply` gibt die Antwort zurück.

**Hinweis:** Rekursive Funktionen müssen einen Basisfall haben, bei dem sie zurückkehren, ohne die Funktion erneut aufzurufen (in diesem Beispiel, wenn `n <= 0`), sonst können sie niemals die Ausführung beenden.

# --instructions--

Schreiben Sie eine rekursive Funktion `sum(arr, n)`, die die Summe der ersten `n` Elemente eines Arrays `arr` zurückgibt.

# --hints--

`sum([1], 0)` sollte 0 sein.

```js
assert.equal(sum([1], 0), 0);
```

`sum([2, 3, 4], 1)` sollte 2 sein.

```js
assert.equal(sum([2, 3, 4], 1), 2);
```

`sum([2, 3, 4, 5], 3)` sollte 9 sein.

```js
assert.equal(sum([2, 3, 4, 5], 3), 9);
```

Ihr Code sollte sich nicht auf irgendeine Art von Schleifen stützen (weder `for` noch `while` oder höherwertige Funktionen wie `forEach`, `map`, `filter` oder `reduce`).

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

Sie sollten Rekursion verwenden, um dieses Problem zu lösen.

```js
assert(
  sum.toString().match(/sum\(.*\)/g).length > 1
);
```

# --seed--

## --seed-contents--

```js
function sum(arr, n) {
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function sum(arr, n) {
  // Only change code below this line
  if(n <= 0) {
    return 0;
  } else {
    return sum(arr, n - 1) + arr[n - 1];
  }
  // Only change code above this line
}
```
