---
id: 5ee127a03c3b35dd45426493
title: Den Wert einer Variablen einer anderen zuweisen
challengeType: 1
forumTopicId: 418265
dashedName: assigning-the-value-of-one-variable-to-another
---

# --description--

Nachdem ein Wert mit dem <dfn>Zuweisungsoperator</dfn> einer Variablen zugewiesen wurde, können Sie den Wert dieser Variablen mit dem <dfn>Zuweisungsoperator</dfn> einer anderen Variablen zuweisen.

```js
var myVar;
myVar = 5;
var myNum;
myNum = myVar;
```

Oben wird eine `myVar`-Variable ohne Wert deklariert und dann der Wert `5` zugewiesen. Anschließend wird eine Variable namens `myNum` ohne Wert deklariert. Dann wird der Inhalt von `myVar` (welcher `5` ist) der Variable `myNum` zugewiesen. Jetzt hat auch `myNum` den Wert von `5`.

# --instructions--

Weisen Sie den Inhalt von `a` der Variablen `b` zu.

# --hints--

Sie sollten den Code oberhalb des angegebenen Kommentars nicht ändern.

```js
assert(/var a;/.test(__helpers.removeJSComments(code)) && /a = 7;/.test(__helpers.removeJSComments(code)) && /var b;/.test(__helpers.removeJSComments(code)));
```

`b` sollte den Wert `7` haben.

```js
assert(typeof b === 'number' && b === 7);
```

`a` sollte mit `b` `=` zugewiesen werden.

```js
assert(/b\s*=\s*a\s*/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
var a;
a = 7;
var b;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
var b;
b = a;
```
