---
id: bd7993c9c69feddfaeb7bdef
title: Zwei Dezimalzahlen mit JavaScript multiplizieren
challengeType: 1
forumTopicId: 301173
dashedName: multiply-two-decimals-with-javascript
---

# --description--

In JavaScript können Sie auch mit Dezimalzahlen rechnen, genauso wie mit ganzen Zahlen.

Multiplizieren wir zwei Dezimalzahlen miteinander, um ihr Produkt zu erhalten.

# --instructions--

Ändern Sie den `0.0`, sodass product gleich `5.0` ist.

# --hints--

Die Variable `product` sollte gleich `5.0` sein.

```js
assert(product === 5.0);
```

Sie sollten den `*`-Operator verwenden.

```js
assert(/\*/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const product = 2.0 * 0.0;
```

# --solutions--

```js
const product = 2.0 * 2.5;
```
