---
id: cf1111c1c11feddfaeb3bdef
title: Zwei Zahlen mit JavaScript addieren
challengeType: 1
forumTopicId: 16650
dashedName: add-two-numbers-with-javascript
---

# --description--

`Number` ist ein Datentyp in JavaScript, der numerische Daten darstellt.

Versuchen wir nun, mit JavaScript zwei Zahlen zu addieren.

JavaScript verwendet das `+`-Symbol als Additionsoperator, wenn es zwischen zwei Zahlen steht.

**Beispiel:**

```js
const myVar = 5 + 10;
```

`myVar` hat jetzt den Wert `15`.

# --instructions--

Ändern Sie das `0`, sodass sum gleich `20` ist.

# --hints--

`sum` sollte gleich `20` sein.

```js
assert(sum === 20);
```

Sie sollten den `+`-Operator verwenden.

```js
assert(/\+/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const sum = 10 + 0;
```

# --solutions--

```js
const sum = 10 + 10;
```
