---
id: cf1231c1c11feddfaeb5bdef
title: Zwei Zahlen mit JavaScript multiplizieren
challengeType: 1
forumTopicId: 18243
dashedName: multiply-two-numbers-with-javascript
---

# --description--

Wir können auch eine Zahl mit einer anderen multiplizieren.

JavaScript verwendet das `*`-Symbol für die Multiplikation von zwei Zahlen.

**Beispiel**

```js
const myVar = 13 * 13;
```

`myVar` hätte den Wert `169`.

# --instructions--

Ändern Sie die `0`, sodass product den Wert `80` hat.

# --hints--

Die Variable `product` sollte den Wert 80 haben.

```js
assert(product === 80);
```

Sie sollten den `*`-Operator verwenden.

```js
assert(/\*/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const product = 8 * 0;
```

# --solutions--

```js
const product = 8 * 10;
```
