---
id: cf1391c1c11feddfaeb4bdef
title: Dezimalzahlen mit JavaScript erstellen
challengeType: 1
forumTopicId: 16826
dashedName: create-decimal-numbers-with-javascript
---

# --description--

Wir können Dezimalzahlen auch in Variablen speichern. Dezimalzahlen werden manchmal als <dfn>Gleitkommazahlen</dfn> oder <dfn>Floats</dfn> bezeichnet.

**Hinweis:** Wenn Sie Zahlen berechnen, erfolgt dies mit endlicher Genauigkeit. Operationen mit Gleitkommazahlen können zu anderen Ergebnissen führen als gewünscht. Wenn Sie eines dieser Ergebnisse erhalten, eröffnen Sie bitte ein Thema im <a href="https://forum.freecodecamp.org/" target="_blank" rel="noopener noreferrer nofollow">freeCodeCamp-Forum</a>.

# --instructions--

Erstellen Sie eine Variable `myDecimal` und weisen Sie ihr einen Dezimalwert mit einem Bruchteil zu (z. B. `5.7`).

# --hints--

`myDecimal` sollte eine Zahl sein.

```js
assert(typeof myDecimal === 'number');
```

`myDecimal` sollte einen Dezimalpunkt enthalten.

```js
assert(myDecimal % 1 != 0);
```

# --seed--

## --seed-contents--

```js
const ourDecimal = 5.7;

// Only change code below this line

```

# --solutions--

```js
const myDecimal = 9.9;
```
