---
id: 587d825a367417b2b2512c89
title: 实现快速排序算法
challengeType: 26
dashedName: implement-the-quicksort-algorithm
---

# --description--

在本实验中你将实现排序算法：quicksort。

快速排序是一种高效的、递归的分治算法，用于对数组进行排序。在此方法中，从原始数组中选择一个枢轴值。然后将数组分区为两个子数组，分别包含小于和大于枢轴值的值。接着递归调用快速排序算法对这两个子数组进行排序，并将结果合并。该过程持续进行，直到达到空数组或单项数组的基本情形，此时返回该数组。递归调用的展开过程返回排序后的数组。

虽然枢轴值的选择很重要，但任何枢轴都可以满足我们的目的。 为简单起见，可以使用第一个或最后一个元素。

Quicksort is a very efficient sorting method, providing *O(nlog(n))* performance on average. 它也相对容易实现。 这些属性使其成为一种流行且有用的排序方法。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个 `quicksort` 函数，它以整数数组作为输入。
1. `quicksort` 函数应返回一个包含与输入相同整数但按从低到高顺序排列的数组。
1. `quicksort` 函数应递归地调用自身以排序数组。

# --hints--

你应该有一个 `quicksort` 函数。

```js
assert.isFunction(quicksort);
```

`quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` 应返回一个除了顺序外未改变的数组。

```js
assert.sameMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quicksort` 应该返回一个排序后的数组（从小到大）。

```js
assert.sameOrderedMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`quicksort` 不应使用内置的 `.sort()` 方法。

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quicksort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function quicksort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quicksort(lesser), ...equal, ...quicksort(greater)];
  }
}
```
