---
id: 66cf315cff4d36b27da828ba
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

이번 단계에서는 플레이어가 라운드에서 이겼는지 판단하는 데 집중하세요.

두 개의 매개변수 `hasPlayerWonTheRound`과 `playerChoice`를 가진 `computerChoice` 함수를 만드세요. 이 매개변수들은 플레이어와 컴퓨터가 선택한 `"Rock"`, `"Paper"`, 또는 `"Scissors"` 중 하나의 문자열 값을 받습니다.

플레이어가 라운드에서 이겼다면 함수는 `true`을 반환하고, 플레이어가 졌거나 비겼다면 `false`을 반환해야 합니다.

플레이어가 라운드에서 이기기 위한 조건은 다음과 같습니다:

- 플레이어가 `"Rock"`를 선택하고 컴퓨터가 `"Scissors"`를 선택한 경우
- 플레이어가 `"Scissors"`를 선택하고 컴퓨터가 `"Paper"`를 선택한 경우
- 플레이어가 `"Paper"`를 선택하고 컴퓨터가 `"Rock"`를 선택한 경우

# --hints--

`hasPlayerWonTheRound`라는 함수를 만들어야 합니다.

```js
assert.isFunction(hasPlayerWonTheRound);
```

함수는 `playerChoice`와 `computerChoice`라는 두 매개변수를 가져야 합니다.

```js
assert.match(code, /hasPlayerWonTheRound\s*(?:=\s*)?\(\s*playerChoice\s*,\s*computerChoice\s*\)\s*/);
```

`hasPlayerWonTheRound` 함수는 불리언을 반환해야 합니다.

```js
assert.isBoolean(hasPlayerWonTheRound("Rock", "Scissors"));
```

`hasPlayerWonTheRound` 함수는 플레이어가 `true`를 선택하고 컴퓨터가 `"Rock"`를 선택했다면 `"Scissors"`를 반환해야 합니다.

```js
assert.isTrue(hasPlayerWonTheRound("Rock", "Scissors"));
```

`hasPlayerWonTheRound` 함수는 플레이어가 `true`를 선택하고 컴퓨터가 `"Scissors"`를 선택했다면 `"Paper"`를 반환해야 합니다.

```js
assert.isTrue(hasPlayerWonTheRound("Scissors", "Paper"));
```

`hasPlayerWonTheRound` 함수는 플레이어가 `true`를 선택하고 컴퓨터가 `"Paper"`를 선택했다면 `"Rock"`를 반환해야 합니다.

```js
assert.isTrue(hasPlayerWonTheRound("Paper", "Rock"));
```

`hasPlayerWonTheRound` 함수는 플레이어와 컴퓨터가 같은 옵션을 선택했다면 `false`를 반환해야 합니다.

```js
assert.isFalse(hasPlayerWonTheRound("Rock", "Rock"));
assert.isFalse(hasPlayerWonTheRound("Scissors", "Scissors"));
assert.isFalse(hasPlayerWonTheRound("Paper", "Paper"));
```

`hasPlayerWonTheRound` 함수는 컴퓨터가 라운드에서 이겼다면 `false`를 반환해야 합니다.

```js
assert.isFalse(hasPlayerWonTheRound("Scissors", "Rock"));
assert.isFalse(hasPlayerWonTheRound("Paper", "Scissors"));
assert.isFalse(hasPlayerWonTheRound("Rock", "Paper"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Rock, Paper, Scissors game</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Let's play Rock, Paper, Scissors!</h1>
    <main>
      <details class="rules-container">
        <summary>Rules to the game</summary>

        <p>You will be playing against the computer.</p>
        <p>You can choose between Rock, Paper, and Scissors.</p>
        <p>The first one to three points wins.</p>

        <p>Here are the rules to getting a point in the game:</p>
        <ul>
          <li>Rock beats Scissors</li>
          <li>Scissors beats Paper</li>
          <li>Paper beats Rock</li>
        </ul>
        <p>
          If the player and computer choose the same option (e.g., Paper and
          Paper), then no one gets the point.
        </p>
      </details>

      <div class="score-container">
        <strong
          >Player Score: <span class="score" id="player-score">0</span></strong
        >
        <strong
          >Computer Score:
          <span class="score" id="computer-score">0</span></strong
        >
      </div>

      <section class="options-container">
        <h2>Choose an option:</h2>
        <div class="btn-container">
          <button id="rock-btn" class="btn">Rock</button>
          <button id="paper-btn" class="btn">Paper</button>
          <button id="scissors-btn" class="btn">Scissors</button>
        </div>
      </section>

      <div class="results-container">
        <p id="results-msg"></p>
        <p id="winner-msg"></p>
        <button class="btn" id="reset-game-btn">Play again?</button>
      </div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --very-dark-blue: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--very-dark-blue);
  text-align: center;
  color: var(--white);
}

h1 {
  margin: 15px 0 20px;
}

.btn {
  cursor: pointer;
  width: 100px;
  margin: 10px;
  color: var(--very-dark-blue);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

.rules-container {
  padding: 10px 0;
  margin: auto;
  border-radius: 15px;
  border: 5px solid var(--yellow);
  background-color: var(--white);
  color: var(--very-dark-blue);
}

.rules-container ul {
  list-style-type: none;
}

.rules-container p {
  margin: 10px 0;
}

@media (min-width: 760px) {
  .rules-container {
    width: 60%;
  }
}

.score-container {
  display: flex;
  justify-content: space-around;
  margin: 30px 0;
  font-size: 1.2rem;
}

.score {
  font-weight: 500;
}

.results-container {
  font-size: 1.3rem;
  margin: 15px 0;
}

#winner-msg {
  margin-top: 25px;
}

#reset-game-btn {
  display: none;
  margin: 20px auto;
}

```

```js
const options = ["Rock", "Paper", "Scissors"];

function getRandomComputerResult() {
  const randomIndex = Math.floor(Math.random() * options.length);
  return options[randomIndex];
}

--fcc-editable-region--

--fcc-editable-region--
```
