---
id: 6140c7e645d8e905819f1dd4
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Para começar, adicione um elemento `link` para o arquivo `./styles.css`.

# --hints--

O código deve ter um elemento `link`.

```js
assert.match(code, /<link/);
```

Você deve ter um elemento `link`.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

O elemento `link` deve estar dentro do elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>

  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css

```
