---
id: 5d822fd413a79914d39e98cc
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

CSS Variables Skyline プロジェクトへようこそ！まず、`link` 要素内に `styles.css` ファイルをリンクする `head` 要素を追加してください。

# --hints--

コードには `link` 要素が必要です。

```js
assert.isNotNull(document.querySelector('link'));
```

既存の `head` タグは変更しないでください。終了タグを削除していないことを確認してください。

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

1つの空要素 `link` が必要です。

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

`link` 要素は `head` 要素の中にある必要があります。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

`link` 要素には、値が `rel` の `stylesheet` 属性が必要です。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link` 要素には、値が `href` の `styles.css` 属性が必要です。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
