---
id: 5d822fd413a79914d39e98d3
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

`div` コンテナ内に4つの `.bb1` 要素をネストしてください。順番に `bb1a`、`bb1b`、`bb1c`、`bb1d` のクラスを付けます。この建物は4つのセクションを持ちます。

# --hints--

4つの新しい `div` 要素を作成してください。

```js
assert.equal(document.querySelectorAll('div')?.length, 6);
```

新しい `div` 要素のうち1つに `bb1a` のクラスを付けてください。

```js
assert.exists(document.querySelector('div.bb1a'));
```

新しい `div` 要素のうち1つに `bb1b` のクラスを付けてください。

```js
assert.exists(document.querySelector('div.bb1b'));
```

新しい `div` 要素のうち1つに `bb1c` のクラスを付けてください。

```js
assert.exists(document.querySelector('div.bb1c'));
```

新しい `div` 要素のうち1つに `bb1d` のクラスを付けてください。

```js
assert.exists(document.querySelector('div.bb1d'));
```

新しい `div` 要素を正しい順序で配置してください。

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert.isTrue(__t('div.bb1a','bb1b') && __t('div.bb1b','bb1c') && __t('div.bb1c','bb1d'));
```

新しい `div` 要素を `.bb1` 要素内に配置してください。

```js
assert.equal(document.querySelectorAll('div.bb1 > div')?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
--fcc-editable-region--
      <div class="bb1"></div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
    
```

