---
id: 5d822fd413a79914d39e98cf
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

`body`（ページ内で最も内側のボックス）が見えますね；その周りのボックスが`html`要素です。`body`に`height`を`100vh`に設定して、ビューポート全体を埋めるようにしてください。`margin`のデフォルトの`body`は、`margin`を`0`に設定して削除します。最後に、何かがビューポートを超えたときに表示されるスクロールバーを隠すために、`overflow`プロパティを`hidden`に設定してください。

# --hints--

`body`セレクターを使うべきです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

`body`には`height`を`100vh`に設定してください。

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

`body`には`margin`を`0`に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

`body`には`overflow`プロパティを`hidden`に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--
    
```

