---
id: 5d822fd413a79914d39e98cd
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

CSSでは、アスタリスクで全てをターゲットにできます。`*` セレクターを使い、`border` の `1px solid black` を与えて、すべてにボーダーを追加してください。これは要素の位置やサイズを視覚化するのに役立つテクニックです。後でこれを削除します。

# --hints--

`*` セレクターを使うべきです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

すべての要素にスタイルを適用するには、`border` プロパティを使うべきです。

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

すべての要素に `1px solid black` のボーダーを付けるべきです。

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
