---
id: 5d822fd413a79914d39e98cc
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

CSS 変数を使ったスカイラインのプロジェクトにようこそ！ Start by adding a `link` element that links your `styles.css` file within the `head` element.

# --hints--

コードには `link` 要素が必要です。

```js
assert.isNotNull(document.querySelector('link'));
```

既存の `head` タグを変更しないでください。 終了タグを削除していないか確認してください。

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

You should have one void `link` element.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

`link` 要素は `head` 要素の中にある必要があります。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

`link` 要素には `rel` 属性があり、値が `stylesheet` に設定されている必要があります。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link` 要素は値が `styles.css` に設定されている `href` 属性を持つ必要があります。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
