---
id: 5d822fd413a79914d39e98da
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

最初の建物はかなり良くできていますね。`div` コンテナの中に新しく3つの `.background-buildings` 要素をネストして、それぞれに `bb2`、`bb3`、`bb4` のクラスを順番に付けてください。これらは背景のためのさらに3つの建物になります。

# --hints--

クラスが `div` の `bb2` を作成してください。

```js
assert.exists(document.querySelector('div.bb2'));
```

クラスが `div` の `bb3` を作成してください。

```js
assert.exists(document.querySelector('div.bb3'));
```

クラスが `div` の `bb4` を作成してください。

```js
assert.exists(document.querySelector('div.bb4'));
```

新しい `div` 要素を3つ作成してください。

```js
assert.lengthOf(document.querySelectorAll('div'), 9);
```

これらの `div` 要素を `.background-buildings` 要素の中に配置してください。

```js
assert.lengthOf(document.querySelector('div.background-buildings')?.children, 4);
```

要素を正しい順番で配置してください。

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert.isTrue(__t('div.bb1','bb2') && __t('div.bb2','bb3') && __t('div.bb3','bb4'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
--fcc-editable-region--

--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}
    
```

