---
id: 587d781e367417b2b2512acb
title: قفل عنصر إلى الحاوية الأصلية باستخدام التمركز المطلق
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLJ7c3'
forumTopicId: 301060
dashedName: lock-an-element-to-its-parent-with-absolute-positioning
---

# --description--

الخيار التالي لخاصية CSS `position` هو `absolute`، الذي يقفل العنصر في مكانه بالنسبة إلى الحاوية الأصلية. على عكس التمركز بـ `relative`، هذا يخرج العنصر من التدفق الطبيعي للوثيقة، فتتجاهله العناصر المحيطة. تُستخدم خصائص إزاحة CSS (top أو bottom و left أو right) لضبط الموضع.

ميزة في التمركز المطلق هي أنه يُقفل بالنسبة إلى أقرب سلف *متمركز* له. إذا نسيت إضافة قاعدة تمركز إلى العنصر الأصل (عادةً باستخدام `position: relative;`)، سيستمر المتصفح في البحث للأعلى في السلسلة وينتهي به الأمر إلى استخدام وسم `body` كافتراضي.

# --instructions--

اقفل العنصر `#searchbar` في أعلى يمين الحاوية الأصلية `section` بتعيين خاصية `position` له إلى `absolute`. أعطه إزاحات `top` و `right` بقيمة 50 بكسل لكل منهما.

# --hints--

يجب أن يكون للعنصر `#searchbar` خاصية `position` مضبوطة على `absolute`.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.position, 'absolute');
```

يجب أن يستخدم شفرتك إزاحة CSS `top` بقيمة 50 بكسل على العنصر `#searchbar`.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.top, '50px');
```

يجب أن يستخدم شفرتك إزاحة CSS `right` بقيمة 50 بكسل على العنصر `#searchbar`.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.right, '50px');
```

# --seed--

## --seed-contents--

```html
<style>
  #searchbar {



  }
  section {
    position: relative;
  }
</style>
<body>
  <h1>Welcome!</h1>
  <section>
    <form id="searchbar">
      <label for="search">Search:</label>
      <input type="search" id="search" name="search">
      <input type="submit" name="submit" value="Go!">
    </form>
  </section>
</body>
```

# --solutions--

```html
<style>
  #searchbar {
    position: absolute;
    top: 50px;
    right: 50px;
  }
  section {
    position: relative;
  }
</style>
<body>
  <h1>Welcome!</h1>
  <section>
    <form id="searchbar">
      <label for="search">Search:</label>
      <input type="search" id="search" name="search">
      <input type="submit" name="submit" value="Go!">
    </form>
  </section>
</body>
```
