---
id: 587d78a6367417b2b2512adc
title: استخدام خاصية CSS Transform skewY لتمديد عنصر على المحور Y
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MZ2uB'
forumTopicId: 301075
dashedName: use-the-css-transform-property-skewy-to-skew-an-element-along-the-y-axis
---

# --description--

بما أن الدالة `skewX()` تميل العنصر المحدد على المحور X بدرجة معينة، فلا عجب أن الخاصية `skewY()` تميل العنصر على المحور Y (الرأسي).

# --instructions--

مِل العنصر الذي يحمل المعرف `top` بمقدار -10 درجات على المحور Y باستخدام خاصية `transform`.

# --hints--

يجب أن يُميل العنصر الذي يحمل المعرف `top` بمقدار -10 درجات على محوره Y.

```js
assert.match(code, /#top\s*?{\s*?.*?\s*?transform:\s*?skewY\(-10deg\);/g);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin: 50px auto;
  }
  #top {
    background-color: red;

  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>

<div id="top"></div>
<div id="bottom"></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin: 50px auto;
  }
  #top {
    background-color: red;
    transform: skewY(-10deg);
  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>
<div id="top"></div>
<div id="bottom"></div>
```
