---
id: 587d78a9367417b2b2512ae8
title: تعلّم كيف تعمل منحنيات بيزيه
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bDrs8'
forumTopicId: 301058
dashedName: learn-how-bezier-curves-work
---

# --description--

التحدٍّ الأخير قدّم خاصية `animation-timing-function` وبعض الكلمات المفتاحية التي تغيّر سرعة الحركة خلال مدتها. يوفر CSS خيارًا آخر غير الكلمات المفتاحية يمنحك تحكمًا أدق في كيفية تنفيذ الحركة، وذلك باستخدام منحنيات بيزيه.

في تحريكات CSS، تُستخدم منحنيات بيزيه مع الدالة `cubic-bezier`. شكل المنحنى يُمثل كيفية تنفيذ الحركة. المنحنى موجود على نظام إحداثيات 1 في 1. المحور السيني لهذا النظام هو مدة الحركة (فكر فيه كمقياس زمني)، والمحور الصادي هو التغير في الحركة.

تتكون دالة `cubic-bezier` من أربع نقاط رئيسية تقع على شبكة 1 في 1: `p0`، `p1`، `p2`، و `p3`. يتم تعيين `p0` و `p3` لك - هما نقطتا البداية والنهاية اللتان تقعان دائمًا عند الأصل (0, 0) و (1, 1) على التوالي. أنت تحدد قيم x و y للنقطتين الأخريين، ومكان وضعهما في الشبكة يحدد شكل المنحنى الذي ستتبعه الحركة. يتم ذلك في CSS من خلال إعلان قيم x و y لنقطتي "وسم الرابط أو وسم الرابط للنص العالي" `p1` و `p2` بالشكل: `(x1, y1, x2, y2)`. لتجميع كل ذلك، إليك مثالًا لمنحنى بيزيه في شفرة CSS:

```css
animation-timing-function: cubic-bezier(0.25, 0.25, 0.75, 0.75);
```

في المثال أعلاه، قيم x و y متساوية لكل نقطة (x1 = 0.25 = y1 و x2 = 0.75 = y2)، وإذا تذكرت من درس الهندسة، ينتج عن ذلك خط يمتد من الأصل إلى النقطة (1, 1). هذه الحركة هي تغير خطي لعنصر خلال مدة الحركة، وهي نفسها استخدام الكلمة المفتاحية `linear`. بعبارة أخرى، تتغير الحركة بسرعة ثابتة.

# --instructions--

للعنصر الذي يحمل المعرف `ball1`، غيّر قيمة خاصية `animation-timing-function` من `linear` إلى القيمة المكافئة لها في دالة `cubic-bezier`. استخدم قيم النقاط المعطاة في المثال أعلاه.

# --hints--

يجب أن تكون قيمة خاصية `animation-timing-function` للعنصر الذي يحمل المعرف `ball1` هي دالة `cubic-bezier` المكافئة للخطية.

```js
const ballOne = document.querySelector('#ball1');
const ballOneStyle = window.getComputedStyle(ballOne); 
assert.equal(
  ballOneStyle?.animationTimingFunction,
    'cubic-bezier(0.25, 0.25, 0.75, 0.75)'
);
```

لا يجب أن تتغير قيمة خاصية `animation-timing-function` للعنصر الذي يحمل المعرف `ball2`.

```js
const ballTwoElement = document.querySelector('#ball2');
const ballTwoStyle = window.getComputedStyle(ballTwoElement); 
const ball2Animation = __helpers.removeWhiteSpace(ballTwoStyle?.animationTimingFunction);
assert.isTrue(
  ball2Animation == 'ease-out' || ball2Animation == 'cubic-bezier(0,0,0.58,1)'
);
```

# --seed--

## --seed-contents--

```html
<style>

  .balls{
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left: 27%;
    animation-timing-function: linear;
  }
  #ball2 {
    left: 56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }

</style>

<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```

# --solutions--

```html
<style>

  .balls{
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left: 27%;
    animation-timing-function: cubic-bezier(0.25, 0.25, 0.75, 0.75);
  }
  #ball2 {
    left: 56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```
