---
id: 587d78a7367417b2b2512ae0
title: استخدام تحريك CSS لتغيير حالة التحويم لزر
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4vZAa'
forumTopicId: 301073
dashedName: use-css-animation-to-change-the-hover-state-of-a-button
---

# --description--

يمكنك استخدام `@keyframes` في CSS لتغيير لون الزر في حالة التحويم.

إليك مثالًا على تغيير عرض صورة عند التحويم:

```html
<style>
  img {
    width: 30px;
  }
  img:hover {
    animation-name: width;
    animation-duration: 500ms;
  }

  @keyframes width {
    100% {
      width: 40px;
    }
  }
</style>

<img src="https://cdn.freecodecamp.org/curriculum/applied-visual-design/google-logo.png" alt="Google's Logo" />
```

# --instructions--

لاحظ أن `ms` تعني ميلي ثانية، حيث 1000ms تساوي 1s.

استخدم `@keyframes` في CSS لتغيير `background-color` لعنصر `button` ليصبح `#4791d0` عندما يحوم المستخدم فوقه. يجب أن تحتوي قاعدة `@keyframes` على إدخال واحد فقط لـ `100%`.

# --hints--

يجب أن تستخدم قاعدة @keyframes اسم التحريك `animation-name` background-color.

```js
assert.match(code, /@keyframes\s+?background-color\s*?{/g);
```

يجب أن تكون هناك قاعدة واحدة تحت `@keyframes` تغير `background-color` إلى `#4791d0` عند 100%.

```js
assert.match(code, /100%\s*?{\s*?background-color:\s*?#4791d0;\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }

  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
  }


</style>

<button>Register</button>
```

# --solutions--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }

  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
  }

  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```
