---
id: 587d78a6367417b2b2512add
title: إنشاء رسم باستخدام CSS
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDWPs6'
forumTopicId: 301048
dashedName: create-a-graphic-using-css
---

# --description--

بالتلاعب بمحدِّدات وخصائص مختلفة، يمكنك صنع أشكال مثيرة للاهتمام. أحد الأسهل لتجربته هو شكل هلال القمر. في هذا التحدٍّ تحتاج إلى العمل مع خاصية `box-shadow` التي تُحدد ظل العنصر، إلى جانب خاصية `border-radius` التي تتحكم في انحناء زوايا العنصر.

ستُنشئ جسمًا دائريًا شفافًا مع ظل واضح مُزاح قليلاً إلى الجانب - الظل سيكون في الواقع شكل القمر الذي تراه.

لإنشاء جسم دائري، يجب تعيين خاصية `border-radius` إلى قيمة 50%.

قد تتذكر من تحدٍّ سابق أن خاصية `box-shadow` تأخذ قيمًا لـ `offset-x`، و`offset-y`، و`blur-radius`، و`spread-radius`، وقيمة `color` بهذا الترتيب. قيمتا `blur-radius` و`spread-radius` اختياريتان.

# --instructions--

تلاعب بالعنصر المربع في المحرر لإنشاء شكل القمر. أولاً، غيّر `background-color` إلى `transparent`، ثم عيّن خاصية `border-radius` إلى 50% لصنع الشكل الدائري. أخيرًا، غيّر خاصية `box-shadow` لتعيين `offset-x` إلى 25px، و`offset-y` إلى 10px، و`blur-radius` إلى 0، و`spread-radius` إلى 0، و`color` إلى `blue`.

# --hints--

يجب تعيين قيمة خاصية `background-color` إلى `transparent`.

```js
assert.match(code,/background-color:\s*?transparent;/gi);
```

يجب تعيين قيمة خاصية `border-radius` إلى `50%`.

```js
assert.match(code,/border-radius:\s*?50%;/gi);
```

يجب تعيين قيمة خاصية `box-shadow` إلى 25px لـ `offset-x`، و10px لـ `offset-y`، و0 لـ `blur-radius`، و0 لـ `spread-radius`، وأخيرًا `blue` لـ `color`.

```js
assert.match(code,/box-shadow:\s*?25px\s+?10px\s+?0(px)?\s+?0(px)?\s+?blue\s*?;/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .center {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 0px;
    box-shadow: 25px 10px 10px 10px green;
  }

</style>
<div class="center"></div>
```

# --solutions--

```html
<style>
  .center {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    width: 100px;
    height: 100px;
    background-color: transparent;
    border-radius: 50%;
    box-shadow: 25px 10px 0 0 blue;
  }
</style>
<div class="center"></div>
```
