---
id: 587d78a9367417b2b2512aea
title: اجعل الحركة أكثر طبيعية باستخدام منحنى بيزيه
challengeType: 0
videoUrl: 'https://scrimba.com/c/c7akWUv'
forumTopicId: 301063
dashedName: make-motion-more-natural-using-a-bezier-curve
---

# --description--

يُحرك هذا التحدٍّ عنصرًا ليُحاكي حركة كرة تُرمى في الهواء. التحديات السابقة تناولت منحنيات بيزيه المكعبة `linear` و `ease-out`، لكن لا أحدهما يصور حركة الرمي بدقة. تحتاج إلى تخصيص منحنى بيزيه لهذا الغرض.

تُكرر خاصية `animation-timing-function` تلقائيًا عند كل إطار رئيسي عندما تُضبط خاصية `animation-iteration-count` على infinite. وبما أن هناك قاعدة إطار رئيسي في منتصف مدة الحركة (عند `50%`)، ينتج عن ذلك تقدمان متطابقان للحركة عند صعود وهبوط الكرة.

منحنى بيزيه المكعب التالي يُحاكي حركة الرمي:

```css
cubic-bezier(0.3, 0.4, 0.5, 1.6);
```

لاحظ أن قيمة y2 أكبر من 1. رغم أن منحنى بيزيه المكعب يُرسم على نظام إحداثيات 1 في 1، ويقبل قيم x من 0 إلى 1 فقط، يمكن تعيين قيمة y لأرقام أكبر من واحد. هذا ينتج حركة ارتداد مثالية لمحاكاة كرة الرمي.

# --instructions--

غيّر قيمة خاصية `animation-timing-function` للعنصر الذي يحمل المعرف `green` إلى دالة `cubic-bezier` بقيم x1, y1, x2, y2 محددة على التوالي بـ 0.311، 0.441، 0.444، 1.649.

# --hints--

يجب أن تكون قيمة خاصية `animation-timing-function` للعنصر الذي يحمل المعرف `green` دالة `cubic-bezier` بقيم x1, y1, x2, y2 كما هو محدد.

```js
const greenElement = document.querySelector('#green');
 const greenStyle = window.getComputedStyle(greenElement);
assert.equal(greenStyle?.animationTimingFunction, 'cubic-bezier(0.311, 0.441, 0.444, 1.649)');
```

# --seed--

## --seed-contents--

```html
<style>
  .balls {
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    top: 60%;
    animation-name: jump;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 25%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 50%;
    animation-timing-function: ease-out;
  }
  #green {
    background: green;
    left: 75%;
    animation-timing-function: cubic-bezier(0.69, 0.1, 1, 0.1);
  }

  @keyframes jump {
    50% {
      top: 10%;
    }
  }
</style>
<div class="balls" id="red"></div>
<div class="balls" id="blue"></div>
<div class="balls" id="green"></div>
```

# --solutions--

```html
<style>
  .balls {
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    top: 60%;
    animation-name: jump;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 25%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 50%;
    animation-timing-function: ease-out;
  }
  #green {
    background: green;
    left: 75%;
    animation-timing-function: cubic-bezier(0.311, 0.441, 0.444, 1.649);
  }

  @keyframes jump {
    50% {
      top: 10%;
    }
  }
</style>
<div class="balls" id="red"></div>
<div class="balls" id="blue"></div>
<div class="balls" id="green"></div>
```
