---
id: 587d78a5367417b2b2512ad6
title: إنشاء تدرج لوني تدريجي باستخدام CSS
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4dpt9'
forumTopicId: 301047
dashedName: create-a-gradual-css-linear-gradient
---

# --description--

تطبيق لون على عناصر HTML لا يقتصر على لون ثابت واحد. توفر CSS إمكانية استخدام انتقالات لونية، تُعرف أيضًا بالتدرجات، على العناصر. يمكن الوصول إلى ذلك من خلال خاصية `background` ودالة `linear-gradient()`. هذه هي الصيغة النحوية العامة:

```css
background: linear-gradient(gradient_direction, color 1, color 2, color 3, ...);
```

المعلمة الأولى تحدد الاتجاه الذي يبدأ منه انتقال اللون - يمكن التعبير عنه بدرجة، حيث يجعل `90deg` التدرج أفقيًا (من اليسار إلى اليمين) و`45deg` يجعل التدرج قطريًا (من أسفل اليسار إلى أعلى اليمين). المعلمات التالية تحدد ترتيب الألوان المستخدمة في التدرج.

مثال:

```css
background: linear-gradient(90deg, red, yellow, rgb(204, 204, 255));
```

# --instructions--

استخدم `linear-gradient()` في `background` لعنصر `div`، واضبطه باتجاه 35 درجة لتغيير اللون من `#CCFFFF` إلى `#FFCCCC`.

# --hints--

يجب أن يحتوي عنصر `div` على `background` من نوع `linear-gradient` بالاتجاه والألوان المحددة.

```js
const divElement = document.querySelector('div');
const divStyle = window.getComputedStyle(divElement); 
assert.match(divStyle?.background, /linear-gradient\(35deg, rgb\(204, 255, 255\), rgb\(255, 204, 204\)\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin: 50px auto;

  }

</style>

<div></div>
```

# --solutions--

```html
<style>
  div {
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin: 50px auto;
    background: linear-gradient(35deg, #CCFFFF, #FFCCCC);
  }
</style>
<div></div>
```
