---
id: 587d78a3367417b2b2512ad0
title: توسيط عنصر أفقيًا باستخدام خاصية الهامش
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLJqU4'
forumTopicId: 301043
dashedName: center-an-element-horizontally-using-the-margin-property
---

# --description--

تقنية أخرى لتحديد الموضع هي توسيط عنصر كتلي أفقيًا. إحدى الطرق لفعل ذلك هي تعيين `margin` بقيمة auto.

تعمل هذه الطريقة مع الصور أيضًا. الصور عناصر مضمَّنة افتراضيًا، لكن يمكن تحويلها إلى عناصر كتلية عند تعيين خاصية `display` إلى `block`.

# --instructions--

وسّط عنصر `div` على الصفحة بإضافة خاصية `margin` بقيمة `auto`.

# --hints--

يجب أن يكون لـ `div` هامش (`margin`) مضبوط على `auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('div')?.margin ,'auto');
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    background-color: blue;
    height: 100px;
    width: 100px;

  }
</style>
<div></div>
```

# --solutions--

```html
<style>
  div {
    background-color: blue;
    height: 100px;
    width: 100px;
    margin: auto;
  }
</style>
<div></div>
```
