---
id: 587d78a4367417b2b2512ad3
title: ضبط لون عناصر مختلفة إلى ألوان تكاملية
challengeType: 0
videoUrl: 'https://scrimba.com/c/cWmPpud'
forumTopicId: 301033
dashedName: adjust-the-color-of-various-elements-to-complementary-colors
---

# --description--

أظهر تحدٍّ الألوان التكاملية أن الألوان المتقابلة على عجلة الألوان تجعل بعضها يبدو أكثر حيوية عند وضعها جنبًا إلى جنب. مع ذلك، قد يكون التباين البصري القوي مزعجًا إذا أُفرط في استخدامه على موقع الويب، وأحيانًا يصعب قراءة النص إذا وُضع على خلفية بلون تكميلي. عمليًا، يكون أحد الألوان هو السائد عادةً، ويُستخدم اللون التكميلي لجذب الانتباه بصريًا إلى محتوى معين في الصفحة.

# --instructions--

ستستخدم هذه الصفحة ظلًا من اللون الأزرق المخضر (`#09A7A1`) كلون سائد، ولونه التكميلي البرتقالي (`#FF790E`) لتسليط الضوء بصريًا على أزرار التسجيل. غيّر `background-color` لكل من `header` و`footer` من الأسود إلى اللون الأزرق المخضر. ثم غيّر `color` نص `h2` إلى الأزرق المخضر أيضًا. أخيرًا، غيّر `background-color` للزر `button` إلى اللون البرتقالي.

# --hints--

يجب أن يحتوي عنصر `header` على `background-color` بقيمة #09A7A1.

```js
const headerElement = document.querySelector('header'); 
const headerStyle = window.getComputedStyle(headerElement);
assert.equal(headerStyle?.backgroundColor, 'rgb(9, 167, 161)');
```

يجب أن يحتوي عنصر `footer` على `background-color` بقيمة #09A7A1.

```js
const footerElement = document.querySelector('footer'); 
const footerStyle = window.getComputedStyle(footerElement);
assert.equal(footerStyle?.backgroundColor, 'rgb(9, 167, 161)');
```

يجب أن يحتوي عنصر `h2` على `color` بقيمة #09A7A1.

```js
const h2Element = document.querySelector('h2'); 
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.color, 'rgb(9, 167, 161)');
```

يجب أن يحتوي عنصر `button` على `background-color` بقيمة #FF790E.

```js
const buttonElement = document.querySelector('button'); 
const buttonStyle = window.getComputedStyle(buttonElement);
assert.equal(buttonStyle?.backgroundColor, 'rgb(255, 121, 14)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: white;
  }
  header {
    background-color: black;
    color: white;
    padding: 0.25em;
  }
  h2 {
    color: black;
  }
  button {
    background-color: white;
  }
  footer {
    background-color: black;
    color: white;
    padding: 0.5em;
  }
</style>
<header>
  <h1>Cooking with FCC!</h1>
</header>
<main>
  <article>
    <h2>Machine Learning in the Kitchen</h2>
    <p>Join this two day workshop that walks through how to implement cutting-edge snack-getting algorithms with a command line interface. Coding usually involves writing exact instructions, but sometimes you need your computer to execute flexible commands, like <code>fetch Pringles</code>.</p>
    <button>Sign Up</button>
  </article>
  <article>
    <h2>Bisection Vegetable Chopping</h2>
    <p>This week-long retreat will level-up your coding ninja skills to actual ninja skills. No longer is the humble bisection search limited to sorted arrays or coding interview questions, applying its concepts in the kitchen will have you chopping carrots in O(log n) time before you know it.</p>
    <button>Sign Up</button>
  </article>
</main>
<br>
<footer>&copy; 2018 FCC Kitchen</footer>
```

# --solutions--

```html
<style>
  body {
    background-color: white;
  }
  header {
    background-color: #09A7A1;
    color: white;
    padding: 0.25em;
  }
  h2 {
    color: #09A7A1;
  }
  button {
    background-color: #FF790E;
  }
  footer {
    background-color: #09A7A1;
    color: white;
    padding: 0.5em;
  }
</style>
<header>
  <h1>Cooking with FCC!</h1>
</header>
<main>
  <article>
    <h2>Machine Learning in the Kitchen</h2>
    <p>Join this two day workshop that walks through how to implement cutting-edge snack-getting algorithms with a command line interface. Coding usually involves writing exact instructions, but sometimes you need your computer to execute flexible commands, like <code>fetch Pringles</code>.</p>
    <button>Sign Up</button>
  </article>
  <article>
    <h2>Bisection Vegetable Chopping</h2>
    <p>This week-long retreat will level-up your coding ninja skills to actual ninja skills. No longer is the humble bisection search limited to sorted arrays or coding interview questions, applying its concepts in the kitchen will have you chopping carrots in O(log n) time before you know it.</p>
    <button>Sign Up</button>
  </article>
</main>
<br>
<footer>&copy; 2018 FCC Kitchen</footer>
```
