---
id: 587d78a4367417b2b2512ad2
title: تعرّف على الألوان الثلاثية
challengeType: 0
forumTopicId: 301057
dashedName: learn-about-tertiary-colors
---

# --description--

شاشات الكمبيوتر والأجهزة تُنتج ألوانًا مختلفة بدمج كميات من الضوء الأحمر والأخضر والأزرق. يُعرف هذا بنموذج الألوان الإضافي RGB في نظرية الألوان الحديثة. الأحمر (R) والأخضر (G) والأزرق (B) تُسمى الألوان الأساسية. خلط لونين أساسيين يُنتج الألوان الثانوية السيان (G + B)، والماجنتا (R + B)، والأصفر (R + G). رأيت هذه الألوان في تحدٍّ الألوان التكاملية. هذه الألوان الثانوية هي تكملة للون الأساسي الذي لم يُستخدم في تكوينها، وتقع مقابلة لذلك اللون الأساسي على عجلة الألوان. على سبيل المثال، الماجنتا مكونة من الأحمر والأزرق، وهي تكملة للون الأخضر.

الألوان الثلاثية هي نتيجة دمج لون أساسي مع أحد ألوانه الثانوية المجاورة. على سبيل المثال، في نموذج ألوان RGB، الأحمر (أساسي) والأصفر (ثانوي) يُنتجان البرتقالي (ثلاثي). هذا يضيف ستة ألوان أخرى إلى عجلة الألوان البسيطة ليصبح المجموع اثني عشر لونًا.

هناك طرق مختلفة لاختيار ألوان تؤدي إلى توليفة متناغمة في التصميم. أحد الأمثلة التي تستخدم الألوان الثلاثية يُسمى نظام الألوان المكملة المنقسمة. يبدأ هذا النظام بلون أساسي، ثم يدمجه مع لونين مجاورين للونه التكميلي. توفر هذه الألوان الثلاثة تباينًا بصريًا قويًا في التصميم، لكنها أكثر رقة من استخدام لونين تكميليين.

إليك ثلاثة ألوان تم إنشاؤها باستخدام نظام المكملة المنقسمة:

<table><thead><tr><th>اللون</th><th>رمز الهيكس</th></tr></thead><thead></thead><tbody><tr><td>البرتقالي</td><td>#FF7F00</td></tr><tr><td>السيان</td><td>#00FFFF</td></tr><tr><td>التوت</td><td>#FF007F</td></tr></tbody></table>

# --instructions--

غيّر خاصية `background-color` في أصناف `orange` و `cyan` و `raspberry` إلى ألوانها المناسبة. تأكد من استخدام رموز الهيكس وليس أسماء الألوان.

# --hints--

عنصر `div` الذي يحمل الصنف `orange` يجب أن تكون خاصية `background-color` له باللون البرتقالي.

```js
const orangeElement = document.querySelector('.orange');
const orangeStyle = window.getComputedStyle(orangeElement);
assert.equal(orangeStyle?.backgroundColor, 'rgb(255, 127, 0)');
```

عنصر `div` الذي يحمل الصنف `cyan` يجب أن تكون خاصية `background-color` له باللون السيان.

```js
const cyanElement = document.querySelector('.cyan');
const cyanStyle = window.getComputedStyle(cyanElement);
assert.equal(cyanStyle?.backgroundColor, 'rgb(0, 255, 255)');
```

عنصر `div` الذي يحمل الصنف `raspberry` يجب أن تكون خاصية `background-color` له باللون التوت.

```js
const raspberryElement = document.querySelector('.raspberry');
const raspberryStyle = window.getComputedStyle(raspberryElement);
assert.equal(raspberryStyle?.backgroundColor, 'rgb(255, 0, 127)');
```

جميع قيم `background-color` لأصناف الألوان يجب أن تكون رموز هيكس وليس أسماء ألوان.

```js
assert.notMatch(code, /background-color:\s(orange|cyan|raspberry)/);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .orange {
    background-color: #000000;
  }

  .cyan {
    background-color: #000000;
  }

  .raspberry {
    background-color: #000000;
  }

  div {
    height: 100px;
    width: 100px;
    margin-bottom: 5px;
  }
</style>

<div class="orange"></div>
<div class="cyan"></div>
<div class="raspberry"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .orange {
    background-color: #FF7F00;
  }

  .cyan {
    background-color: #00FFFF;
  }

  .raspberry {
    background-color: #FF007F;
  }

  div {
    height: 100px;
    width: 100px;
    margin-bottom: 5px;
  }
</style>
<div class="orange"></div>
<div class="cyan"></div>
<div class="raspberry"></div>
```
