---
id: 587d781a367417b2b2512ab7
title: استخدام وسم strong لجعل النص عريضًا
challengeType: 0
videoUrl: 'https://scrimba.com/c/ceJNBSb'
forumTopicId: 301080
dashedName: use-the-strong-tag-to-make-text-bold
---

# --description--

لجعل النص عريضًا، يمكنك استخدام وسم `strong`. يُستخدم هذا غالبًا لجذب الانتباه إلى النص وللدلالة على أهميته. مع وسم `strong`، يطبق المتصفح خاصية CSS `font-weight: bold;` على العنصر.

# --instructions--

لفّ وسم `strong` حول النص `Stanford University` داخل وسم `p` (لا تشمل النقطة).

# --hints--

يجب أن يضيف شفرتك وسم `strong` واحد إلى الترميز.

```js
assert.lengthOf(document.querySelectorAll('strong'),1);
```

يجب أن يكون وسم `strong` داخل وسم `p`.

```js
const paragraphElement = document.querySelector('p');
const paragraphChildren = paragraphElement?.querySelectorAll(`:scope ${'strong'}`);
assert.lengthOf(paragraphChildren, 1);
```

يجب أن يلف وسم `strong` حول الكلمات `Stanford University`.

```js
assert.match(document.querySelector('strong')?.textContent, /^Stanford University\.?$/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at Stanford University.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at <strong>Stanford University</strong>.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
