---
id: 587d78a5367417b2b2512ad7
title: استخدام تدرج خطي في CSS لإنشاء عنصر مخطط
challengeType: 0
videoUrl: 'https://scrimba.com/c/c6bmQh2'
forumTopicId: 301072
dashedName: use-a-css-linear-gradient-to-create-a-striped-element
---

# --description--

الدالة `repeating-linear-gradient()` تشبه كثيرًا الدالة `linear-gradient()` مع الفرق الرئيسي أنها تكرر نمط التدرج المحدد. تقبل `repeating-linear-gradient()` مجموعة متنوعة من القيم، لكن لتبسيط الأمر، ستعمل في هذا التحدٍّ مع قيمة زاوية وقيم توقفات اللون.

قيمة الزاوية هي اتجاه التدرج. توقفات اللون تشبه قيم العرض التي تحدد مكان حدوث الانتقال، وتُعطى بنسبة مئوية أو بعدد من البيكسلات.

في المثال المعروض في محرر الشفرة، يبدأ التدرج باللون `yellow` عند 0 بيكسل، والذي يمتزج مع اللون الثاني `blue` على بعد 40 بيكسل من البداية. وبما أن توقف اللون التالي أيضًا عند 40 بيكسل، يتغير التدرج فورًا إلى اللون الثالث `green`، والذي يمتزج بدوره مع اللون الرابع `red` عند 80 بيكسل من بداية التدرج.

لهذا المثال، من المفيد التفكير في توقفات اللون كأزواج حيث يمتزج كل لونين معًا.

```css
0px [yellow -- blend -- blue] 40px [green -- blend -- red] 80px
```

إذا كانت كل قيم توقف اللونين متطابقة، فإن الامتزاج لا يكون ملحوظًا لأنه بين نفس اللون، يتبعه انتقال حاد إلى اللون التالي، فتنتج في النهاية خطوطًا مخططة.

# --instructions--

اصنع خطوطًا مخططة بتغيير `repeating-linear-gradient()` لاستخدام زاوية تدرج `45deg`، ثم عيّن أول توقفين للون إلى `yellow`، وأخيرًا عيّن التوقفين الثانيين إلى `black`.

# --hints--

يجب أن تكون زاوية `repeating-linear-gradient()` 45deg.

```js
assert.match(code,/background:\s*?repeating-linear-gradient\(\s*?45deg/gi);
```

يجب ألا تكون زاوية `repeating-linear-gradient()` 90deg بعد الآن.

```js
assert.notMatch(code, /90deg/gi);
```

يجب أن يكون توقف اللون عند 0 بيكسل هو `yellow`.

```js
assert.match(code, /yellow\s+?0(px)?/gi);
```

يجب أن يكون أول توقف للون عند 40 بيكسل هو `yellow`.

```js
assert.match(code, /yellow\s+?40px/gi);
```

يجب أن يكون التوقف الثاني للون عند 40 بيكسل هو `black`.

```js
assert.match(code, /yellow\s+?40px,\s*?black\s+?40px/gi);
```

يجب أن يكون آخر توقف للون عند 80 بيكسل هو `black`.

```js
assert.match(code, /black\s+?80px/gi);
```

# --seed--

## --seed-contents--

```html
<style>

  div{
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin:  50 auto;
    background: repeating-linear-gradient(
      90deg,
      yellow 0px,
      blue 40px,
      green 40px,
      red 80px
    );
  }

</style>

<div></div>
```

# --solutions--

```html
<style>
  div{
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin:  50 auto;
    background: repeating-linear-gradient(
      45deg,
      yellow 0px,
      yellow 40px,
      black 40px,
      black 80px
    );
  }
</style>
<div></div>
```
