---
id: 587d78a7367417b2b2512ae1
title: إنشاء حركة باستخدام تحريك CSS
challengeType: 0
videoUrl: 'https://scrimba.com/c/c7amZfW'
forumTopicId: 301051
dashedName: create-movement-using-css-animation
---

# --description--

عندما تحدد عناصر خاصية `position` مثل `fixed` أو `relative`، يمكن استخدام خصائص إزاحة CSS `right` و`left` و`top` و`bottom` في قواعد التحريك لإنشاء حركة.

كما هو موضح في المثال أدناه، يمكنك دفع العنصر للأسفل ثم للأعلى بتعيين خاصية `top` في الإطار المفتاحي `50%` إلى 50px، مع إبقائها على 0px في الإطارين المفتاحيين الأول (`0%`) والأخير (`100%`).

```css
@keyframes rainbow {
  0% {
    background-color: blue;
    top: 0px;
  }
  50% {
    background-color: green;
    top: 50px;
  }
  100% {
    background-color: yellow;
    top: 0px;
  }
}
```

# --instructions--

أضف حركة أفقية إلى تحريك `div`. باستخدام خاصية إزاحة `left`، أضف إلى قاعدة `@keyframes` بحيث يبدأ قوس قزح عند 0 بكسل عند `0%`، ويتحرك إلى 25 بكسل عند `50%`، وينتهي عند -25 بكسل عند `100%`. لا تستبدل خاصية `top` في المحرر — يجب أن يحتوي التحريك على حركة عمودية وأفقية معًا.

# --hints--

يجب أن تستخدم قاعدة `@keyframes` عند `0%` إزاحة `left` بقيمة 0px.

```js
assert.match(code,/[^50]0%\s*?{[\s\S]*?left:\s*?0px(;[\s\S]*?|\s*?)}/gi);
```

يجب أن تستخدم قاعدة `@keyframes` عند `50%` إزاحة `left` بقيمة 25px.

```js
assert.match(code,/50%\s*?{[\s\S]*?left:\s*?25px(;[\s\S]*?|\s*?)}/gi);
```

يجب أن تستخدم قاعدة `@keyframes` عند `100%` إزاحة `left` بقيمة -25px.

```js
assert.match(code,/100%\s*?{[\s\S]*?left:\s*?-25px(;[\s\S]*?|\s*?)}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
    position: relative;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
      top: 0px;

    }
    50% {
      background-color: green;
      top: 50px;

    }
    100% {
      background-color: yellow;
      top: 0px;

    }
  }
</style>

<div id="rect"></div>
```

# --solutions--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
    position: relative;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
      top: 0px;
      left: 0px;
    }
    50% {
      background-color: green;
      top: 50px;
      left: 25px;
    }
    100% {
      background-color: yellow;
      top: 0px;
      left: -25px;
    }
  }
</style>
<div id="rect"></div>
```
