---
id: 587d78a7367417b2b2512adf
title: تعرّف كيف تعمل خاصيات CSS @keyframes و animation
challengeType: 0
videoUrl: 'https://scrimba.com/c/cakprhv'
forumTopicId: 301059
dashedName: learn-how-the-css-keyframes-and-animation-properties-work
---

# --description--

لتحريك عنصر، تحتاج إلى معرفة خصائص التحريك وقاعدة `@keyframes`. تتحكم خصائص التحريك في كيفية تصرف التحريك، وقاعدة `@keyframes` تتحكم فيما يحدث أثناء هذا التحريك. هناك ثماني خصائص تحريك إجمالًا. سيبقي هذا التحدٍّ الأمر بسيطًا ويغطي أهم خاصيتين أولًا:

تحدد `animation-name` اسم التحريك، الذي يستخدم لاحقًا بواسطة `@keyframes` لإخبار CSS بالقواعد المرتبطة بكل تحريك.

تحدد `animation-duration` مدة التحريك الزمنية.

تُستخدم `@keyframes` لتحديد ما يحدث بالضبط داخل التحريك خلال مدته. يتم ذلك بإعطاء خصائص CSS لإطارات محددة خلال التحريك، مع نسب مئوية تتراوح بين 0% و100%. إذا قارنت هذا بفيلم، فإن خصائص CSS عند 0% تمثل كيف يظهر العنصر في المشهد الافتتاحي. وخصائص CSS عند 100% تمثل كيف يظهر العنصر في النهاية، قبل عرض الاعتمادات مباشرة. ثم تطبق CSS السحر لتحويل العنصر خلال المدة المحددة ليؤدي المشهد. إليك مثالًا لتوضيح استخدام `@keyframes` وخصائص التحريك:

```css
#anim {
  animation-name: colorful;
  animation-duration: 3s;
}

@keyframes colorful {
  0% {
    background-color: blue;
  }
  100% {
    background-color: yellow;
  }
}
```

بالنسبة للعنصر الذي يحمل المعرف `anim`، يحدد مقتطف الشفرة أعلاه `animation-name` إلى `colorful` ويحدد `animation-duration` إلى 3 ثوانٍ. ثم تربط قاعدة `@keyframes` خصائص التحريك بالاسم `colorful`. تحدد اللون إلى الأزرق في بداية التحريك (0%) والذي سينتقل إلى الأصفر في نهاية التحريك (100%). لست مقيدًا بالانتقالات من البداية إلى النهاية فقط، يمكنك تعيين خصائص للعنصر لأي نسبة مئوية بين 0% و100%.

# --instructions--

أنشئ تحريكًا للعنصر الذي يحمل المعرف `rect`، بتعيين `animation-name` إلى `rainbow` و`animation-duration` إلى 4 ثوانٍ. بعد ذلك، أعلن قاعدة `@keyframes`، وحدد `background-color` في بداية التحريك (`0%`) إلى الأزرق، وفي منتصف التحريك (`50%`) إلى الأخضر، وفي نهاية التحريك (`100%`) إلى الأصفر.

# --hints--

يجب أن يحتوي العنصر الذي يحمل المعرف `rect` على خاصية `animation-name` بقيمة `rainbow`.

```js
const rectElement = document.querySelector('#rect');
const rectStyle = window.getComputedStyle(rectElement);
assert.equal(rectStyle?.animationName, 'rainbow');
```

يجب أن يحتوي العنصر الذي يحمل المعرف `rect` على خاصية `animation-duration` بقيمة 4s.

```js
const rectElement = document.querySelector('#rect');
const rectStyle = window.getComputedStyle(rectElement);
assert.equal(rectStyle?.animationDuration, '4s');
```

يجب أن تستخدم قاعدة `@keyframes` اسم `animation-name` بقيمة `rainbow`.

```js
assert.match(code ,/@keyframes\s+?rainbow\s*?{/g);
```

يجب أن تستخدم قاعدة `@keyframes` الخاصة بـ `rainbow` خاصية `background-color` بقيمة `blue` عند 0%.

```js
assert.match(code,/0%\s*?{\s*?background-color:\s*?blue;\s*?}/gi);
```

يجب أن تستخدم قاعدة `@keyframes` الخاصة بـ `rainbow` خاصية `background-color` بقيمة `green` عند 50%.

```js
assert.match(code,/50%\s*?{\s*?background-color:\s*?green;\s*?}/gi);
```

يجب أن تستخدم قاعدة `@keyframes` الخاصة بـ rainbow خاصية `background-color` بقيمة `yellow` عند 100%.

```js
assert.match(code,/100%\s*?{\s*?background-color:\s*?yellow;\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
  }

  #rect {


  }




</style>
<div id="rect"></div>
```

# --solutions--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
    }
    50% {
      background-color: green;
    }
    100% {
      background-color: yellow;
    }
  }
</style>
<div id="rect"></div>
```
