---
id: 587d78a4367417b2b2512ad5
title: ضبط درجة لون
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDJvT7'
forumTopicId: 301038
dashedName: adjust-the-tone-of-a-color
---

# --description--

تُسهل خاصية `hsl()` في CSS ضبط درجة اللون أيضًا. خلط الأبيض مع لون نقي يُنتج ظلًا فاتحًا لذلك اللون، وإضافة الأسود تُنتج ظلًا داكنًا. بدلاً من ذلك، تُنتج الدرجة بإضافة الرمادي أو عن طريق التلوين والتظليل معًا. تذكر أن الحرفين 's' و'l' في `hsl()` يرمزان إلى التشبع ودرجة الفاتح على التوالي. نسبة التشبع تُغير كمية الرمادي، ونسبة الفاتح تحدد مقدار الأبيض أو الأسود في اللون. هذا مفيد عندما يكون لديك لون أساسي تحبه، لكن تحتاج إلى تباينات مختلفة منه.

# --instructions--

جميع العناصر لها `background-color` افتراضي بقيمة `transparent`. يظهر عنصر `nav` لدينا حاليًا بخلفية `cyan`، لأن العنصر خلفه له `background-color` مضبوط على `cyan`. أضف `background-color` إلى عنصر `nav` ليستخدم نفس درجة اللون `cyan`، لكن بقيم تشبع `80%` وفاتح اللون `25%` لتغيير نغمة وظل اللون.

# --hints--

يجب أن يكون لعنصر `nav` خاصية `background-color` بدرجة اللون الأزرق المخضر المعدلة باستخدام خاصية `hsl()`.

```js
// Computed style of hsl(180, 80%, 25%) results in rgb(13,115,115)
assert.equal(
  new __helpers.CSSHelp(document).getStyle('nav').getPropVal('background-color', true), 
  'rgb(13,115,115)'
)
```

# --seed--

## --seed-contents--

```html
<style>
  header {
    background-color: hsl(180, 90%, 35%);
    color: #FFFFFF;
  }

  nav {

  }

  h1 {
    text-indent: 10px;
    padding-top: 10px;
  }

  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }

  nav li {
    display: inline;
    margin-right: 20px;
  }

  a {
    text-decoration: none;
    color: inherit;
  }
</style>

<header>
  <h1>Cooking with FCC!</h1>
  <nav>
    <ul>
      <li><a href="#">Home</a></li>
      <li><a href="#">Classes</a></li>
      <li><a href="#">Contact</a></li>
    </ul>
  </nav>
</header>
```

# --solutions--

```html
<style>
  header {
    background-color: hsl(180, 90%, 35%);
    color: #FFFFFF;
  }

  nav {
    background-color: hsl(180, 80%, 25%);
  }

  h1 {
    text-indent: 10px;
    padding-top: 10px;
  }

  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }

  nav li {
    display: inline;
    margin-right: 20px;
  }

  a {
    text-decoration: none;
    color: inherit;
  }
</style>
<header>
  <h1>Cooking with FCC!</h1>
  <nav>
    <ul>
      <li><a href="#">Home</a></li>
      <li><a href="#">Classes</a></li>
      <li><a href="#">Contact</a></li>
    </ul>
  </nav>
</header>
```
