---
id: 587d78a5367417b2b2512ad8
title: إنشاء نسيج بإضافة نمط خفيف كصورة خلفية
challengeType: 0
videoUrl: 'https://scrimba.com/c/cQdwJC8'
forumTopicId: 301052
dashedName: create-texture-by-adding-a-subtle-pattern-as-a-background-image
---

# --description--

إحدى الطرق لإضافة نسيج وجاذبية إلى الخلفية وجعلها تبرز أكثر هي إضافة نمط خفيف. المفتاح هو التوازن، إذ لا تريد أن تبرز الخلفية كثيرًا فتأخذ الانتباه بعيدًا عن المقدمة. خاصية `background` تدعم الدالة `url()` لربط صورة النسيج أو النمط المختار. عنوان الرابط يُحاط بعلامات اقتباس داخل الأقواس العادية.

# --instructions--

باستخدام الرابط `https://cdn-media-1.freecodecamp.org/imgr/MJAkxbh.png`، عيّن خاصية `background` للصفحة كاملة باستخدام محدد `body`.

# --hints--

يجب أن يحتوي عنصر `body` على خاصية `background` مضبوطة إلى دالة `url()` مع الرابط المعطى.

```js
assert.match(code,/background(-image)?:\s*?url\(\s*("|'|)https:\/\/cdn-media-1\.freecodecamp\.org\/imgr\/MJAkxbh\.png\2\s*\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {

  }
</style>
```

# --solutions--

```html
<style>
  body {
    background: url("https://cdn-media-1.freecodecamp.org/imgr/MJAkxbh.png");
  }
</style>
```
