---
id: 587d78a8367417b2b2512ae5
title: تحريك العناصر بمعدلات متغيرة
challengeType: 0
videoUrl: 'https://scrimba.com/c/cZ89WA4'
forumTopicId: 301040
dashedName: animate-elements-at-variable-rates
---

# --description--

هناك عدة طرق لتغيير معدلات التحريك للعناصر المتحركة بشكل مشابه. حتى الآن، تم تحقيق ذلك بتطبيق خاصية `animation-iteration-count` وتعيين قواعد `@keyframes`.

لتوضيح ذلك، يتكون التحريك على اليمين من نجمتين كل واحدة منهما تقلص في الحجم والشفافية عند علامة 20% في قاعدة `@keyframes`، مما يخلق تأثير الوميض. يمكنك تغيير قاعدة `@keyframes` لأحد العناصر بحيث تومض النجوم بمعدلات مختلفة.

# --instructions--

غيّر معدل التحريك للعنصر الذي يحمل اسم الصنف `star-1` بتغيير قاعدة `@keyframes` الخاصة به إلى 50%.

# --hints--

يجب أن تكون قاعدة `@keyframes` لصنف `star-1` عند 50%.

```js
assert.match(code,/twinkle-1\s*?{\s*?50%/g);
```

# --seed--

## --seed-contents--

```html
<style>
  .stars {
    background-color: white;
    height: 30px;
    width: 30px;
    border-radius: 50%;
    animation-iteration-count: infinite;
  }

  .star-1 {
    margin-top: 15%;
    margin-left: 60%;
    animation-name: twinkle-1;
    animation-duration: 1s;
  }

  .star-2 {
    margin-top: 25%;
    margin-left: 25%;
    animation-name: twinkle-2;
    animation-duration: 1s;
  }

  @keyframes twinkle-1 {
    20% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  @keyframes twinkle-2 {
    20% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  #back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: linear-gradient(black, #000099, #66c2ff, #ffcccc, #ffeee6);
  }
</style>

<div id="back"></div>
<div class="star-1 stars"></div>
<div class="star-2 stars"></div>
```

# --solutions--

```html
<style>
  .stars {
    background-color: white;
    height: 30px;
    width: 30px;
    border-radius: 50%;
    animation-iteration-count: infinite;
  }

  .star-1 {
    margin-top: 15%;
    margin-left: 60%;
    animation-name: twinkle-1;
    animation-duration: 1s;
  }

  .star-2 {
    margin-top: 25%;
    margin-left: 25%;
    animation-name: twinkle-2;
    animation-duration: 1s;
  }

  @keyframes twinkle-1 {
    50% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  @keyframes twinkle-2 {
    20% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  #back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: linear-gradient(black, #000099, #66c2ff, #ffcccc, #ffeee6);
  }
</style>
<div id="back"></div>
<div class="star-1 stars"></div>
<div class="star-2 stars"></div>
```
