---
id: 587d78a8367417b2b2512ae6
title: تحريك عدة عناصر بمعدلات متغيرة
challengeType: 0
videoUrl: 'https://scrimba.com/c/cnpWZc9'
forumTopicId: 301042
dashedName: animate-multiple-elements-at-variable-rates
---

# --description--

في التحدٍّ السابق، غيرت معدلات التحريك لعنصرين متحركين بطريقة مشابهة عبر تعديل قواعد `@keyframes`. يمكنك تحقيق نفس الهدف من خلال تعديل `animation-duration` لعدة عناصر.

في التحريك الجاري في محرر الشفرة، هناك ثلاث نجوم في السماء تلمع بنفس المعدل في حلقة مستمرة. لجعلها تلمع بمعدلات مختلفة، يمكنك ضبط خاصية `animation-duration` بقيم مختلفة لكل عنصر.

# --instructions--

اضبط `animation-duration` للعناصر التي تحمل الأصناف `star-1` و`star-2` و`star-3` إلى 1s و0.9s و1.1s على التوالي.

# --hints--

يجب أن تبقى خاصية `animation-duration` للنجم الذي يحمل الصنف `star-1` عند 1s.

```js
const starOne = document.querySelector('.star-1');
const starOneStyle = window.getComputedStyle(starOne); 
assert.equal(starOneStyle?.animationDuration, '1s');
```

يجب أن تكون خاصية `animation-duration` للنجم الذي يحمل الصنف `star-2` 0.9s.

```js
const starTwo = document.querySelector('.star-2');
const starTwoStyle = window.getComputedStyle(starTwo); 
assert.equal(starTwoStyle?.animationDuration , '0.9s');
```

يجب أن تكون خاصية `animation-duration` للنجم الذي يحمل الصنف `star-3` 1.1s.

```js
const starThree = document.querySelector('.star-3');
const starThreeStyle = window.getComputedStyle(starThree); 
assert.equal(starThreeStyle?.animationDuration, '1.1s');
```

# --seed--

## --seed-contents--

```html
<style>
  .stars {
    background-color: white;
    height: 30px;
    width: 30px;
    border-radius: 50%;
    animation-iteration-count: infinite;
  }

  .star-1 {
    margin-top: 15%;
    margin-left: 60%;
    animation-duration: 1s;
    animation-name: twinkle;
  }

  .star-2 {
    margin-top: 25%;
    margin-left: 25%;
    animation-duration: 1s;
    animation-name: twinkle;
  }

  .star-3 {
    margin-top: 10%;
    margin-left: 50%;
    animation-duration: 1s;
    animation-name: twinkle;
  }

  @keyframes twinkle {
    20% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  #back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: linear-gradient(black, #000099, #66c2ff, #ffcccc, #ffeee6);
  }
</style>

<div id="back"></div>
<div class="star-1 stars"></div>
<div class="star-2 stars"></div>
<div class="star-3 stars"></div>
```

# --solutions--

```html
<style>
  .stars {
    background-color: white;
    height: 30px;
    width: 30px;
    border-radius: 50%;
    animation-iteration-count: infinite;
  }

  .star-1 {
    margin-top: 15%;
    margin-left: 60%;
    animation-duration: 1s;
    animation-name: twinkle;
  }

  .star-2 {
    margin-top: 25%;
    margin-left: 25%;
    animation-duration: 0.9s;
    animation-name: twinkle;
  }

  .star-3 {
    margin-top: 10%;
    margin-left: 50%;
    animation-duration: 1.1s;
    animation-name: twinkle;
  }

  @keyframes twinkle {
    20% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  #back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: linear-gradient(black, #000099, #66c2ff, #ffcccc, #ffeee6);
  }
</style>
<div id="back"></div>
<div class="star-1 stars"></div>
<div class="star-2 stars"></div>
<div class="star-3 stars"></div>
```
