---
id: 587d78a6367417b2b2512ade
title: إنشاء شكل أكثر تعقيدًا باستخدام CSS و HTML
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPpz4fr'
forumTopicId: 301050
dashedName: create-a-more-complex-shape-using-css-and-html
---

# --description--

أحد أكثر الأشكال شهرة في العالم هو شكل القلب، وفي هذا التحدٍّ ستنشئ واحدًا باستخدام CSS نقي. لكن أولًا، عليك فهم العنصرين المزيفين `::before` و `::after`. ينشئ `::before` عنصرًا مزيفًا يكون الطفل الأول للعنصر المحدد؛ بينما ينشئ `::after` عنصرًا مزيفًا يكون الطفل الأخير للعنصر المحدد. في المثال التالي، يُستخدم عنصر مزيف `::before` لإضافة مستطيل إلى عنصر يحمل الصنف `heart`:

```css
.heart::before {
  content: "";
  background-color: yellow;
  border-radius: 25%;
  position: absolute;
  height: 50px;
  width: 70px;
  top: -50px;
  left: 5px;
}
```

لكي يعمل العنصران المزيفان `::before` و`::after` بشكل صحيح، يجب أن يحتوي كل منهما على خاصية `content` معرفة. تُستخدم هذه الخاصية عادةً لإضافة أشياء مثل صورة أو نص إلى العنصر المحدد. عندما يُستخدم العنصران المزيفان `::before` و`::after` لصنع أشكال، تبقى خاصية `content` مطلوبة، لكنها تُضبط على سلسلة نصية فارغة. في المثال أعلاه، يحتوي العنصر الذي يحمل الصنف `heart` على عنصر مزيف `::before` يُنتج مستطيلاً أصفر بارتفاع وعرض يبلغان `50px` و`70px` على التوالي. هذا المستطيل له زوايا مستديرة بسبب انحناء الزوايا `border-radius` بنسبة 25%، ويُوضع بشكل مطلق على بعد `5px` من اليسار و`50px` فوق الجزء العلوي من العنصر.

# --instructions--

حوّل العنصر على الشاشة إلى شكل قلب. في المحدِّد `.heart::after`، غيّر `background-color` إلى `pink` و`border-radius` إلى 50%.

بعد ذلك، استهدف العنصر الذي يحمل الصنف `heart` (فقط `heart`) واملأ خاصية `transform`. استخدم الدالة `rotate()` مع زاوية -45 درجة.

أخيرًا، في المحدد `.heart::before`، اضبط خاصية `content` على سلسلة نصية فارغة.

# --hints--

يجب أن تكون خاصية `background-color` في المحدِّد `.heart::after` هي `pink`.

```js
const heartAfter = code.match(/\.heart::after\s*{[\s\S]+?[^\}]}/g)?.[0];
assert.match(heartAfter, /({|;)\s*background-color\s*:\s*pink\s*(;|})/g);
```

يجب أن يكون انحناء الزوايا `border-radius` في المحدد `.heart::after` بنسبة 50%.

```js
assert.lengthOf(code.match(/border-radius\s*?:\s*?50%/gi),2);
```

يجب أن تستخدم خاصية `transform` للصنف `heart` دالة `rotate()` مضبوطة على -45 درجة.

```js
assert.match(code,/transform\s*?:\s*?rotate\(\s*?-45deg\s*?\)/gi);
```

يجب أن تكون خاصية `content` في المحدد `.heart::before` سلسلة نصية فارغة.

```js
assert.match(code,/\.heart::before\s*?{\s*?content\s*?:\s*?("|')\1\s*?;/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: ;
  }
  .heart::after {
    background-color: blue;
    content: "";
    border-radius: 25%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart::before {
    content: ;
    background-color: pink;
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }
</style>
<div class="heart"></div>
```

# --solutions--

```html
<style>
  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: rotate(-45deg);
  }
  .heart::after {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart::before {
    content: "";
    background-color: pink;
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }
</style>
<div class="heart"></div>
```
