---
id: 587d78a8367417b2b2512ae3
title: تحريك العناصر باستمرار باستخدام عدد تكرار حركة لا نهائي
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJDVfq'
forumTopicId: 301041
dashedName: animate-elements-continually-using-an-infinite-animation-count
---

# --description--

التحديات السابقة تناولت كيفية استخدام بعض خصائص الحركة وقاعدة `@keyframes`. خاصية حركة أخرى هي `animation-iteration-count`، التي تسمح لك بالتحكم في عدد المرات التي تريد تكرار الحركة خلالها. إليك مثالًا:

```css
animation-iteration-count: 3;
```

في هذه الحالة ستتوقف الحركة بعد تشغيلها 3 مرات، لكن يمكن جعل الحركة تستمر بلا توقف بتعيين هذه القيمة إلى `infinite`.

# --instructions--

للحفاظ على ارتداد الكرة على اليمين في حلقة مستمرة، غيّر خاصية `animation-iteration-count` إلى `infinite`.

# --hints--

يجب أن تكون قيمة خاصية `animation-iteration-count` هي `infinite`.

```js
const ballElement = document.querySelector('#ball');
const ballStyle = window.getComputedStyle(ballElement); 
assert.equal(ballStyle?.animationIterationCount, 'infinite');
```

# --seed--

## --seed-contents--

```html
<style>

  #ball {
    width: 100px;
    height: 100px;
    margin: 50px auto;
    position: relative;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: bounce;
    animation-duration: 1s;
    animation-iteration-count: 3;
  }

  @keyframes bounce{
    0% {
      top: 0px;
    }
    50% {
      top: 249px;
      width: 130px;
      height: 70px;
    }
    100% {
      top: 0px;
    }
  }
</style>
<div id="ball"></div>
```

# --solutions--

```html
<style>
  #ball {
    width: 100px;
    height: 100px;
    margin: 50px auto;
    position: relative;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: bounce;
    animation-duration: 1s;
    animation-iteration-count: infinite;
  }

  @keyframes bounce{
    0% {
      top: 0px;
    }
    50% {
      top: 249px;
      width: 130px;
      height: 70px;
    }
    100% {
      top: 0px;
    }
  }
</style>
<div id="ball"></div>
```
