---
id: 587d78a9367417b2b2512ae9
title: استخدم منحنى بيزيه لتحريك رسم بياني
challengeType: 0
videoUrl: 'https://scrimba.com/c/c6bnRCK'
forumTopicId: 301071
dashedName: use-a-bezier-curve-to-move-a-graphic
---

# --description--

ناقش تحدٍّ سابق الكلمة المفتاحية `ease-out` التي تصف تغيرًا في الحركة يتسارع أولًا ثم يبطئ في نهاية الحركة. على اليمين، يُوضح الفرق بين الكلمة المفتاحية `ease-out` (للعنصر الأزرق) والكلمة المفتاحية `linear` (للعنصر الأحمر). يمكن تحقيق تقدمات حركة مشابهة لـ `ease-out` باستخدام دالة منحنى بيزيه مكعبة مخصصة.

بشكل عام، تغيير نقاط وسم الرابط `p1` و `p2` يُنتج منحنيات بيزيه مختلفة، تتحكم في كيفية تقدم الحركة عبر الزمن. إليك مثالًا لمنحنى بيزيه يستخدم قيمًا لمحاكاة نمط `ease-out`:

```css
animation-timing-function: cubic-bezier(0, 0, 0.58, 1);
```

تذكر أن جميع دوال `cubic-bezier` تبدأ بـ `p0` عند (0, 0) وتنتهي بـ `p3` عند (1, 1). في هذا المثال، يتحرك المنحنى أسرع عبر المحور Y (يبدأ من 0، ثم يصل إلى قيمة y لـ `p1` وهي 0، ثم إلى قيمة y لـ `p2` وهي 1) مقارنةً بحركته عبر المحور X (يبدأ من 0، ثم 0 لـ `p1`، ويصل إلى 0.58 لـ `p2`). نتيجة لذلك، يتقدم التغير في العنصر المتحرك أسرع من زمن الحركة لذلك الجزء. قرب نهاية المنحنى، تنعكس العلاقة بين التغير في قيم x و y — تتحرك قيمة y من 1 إلى 1 (بدون تغير)، بينما تتحرك قيم x من 0.58 إلى 1، مما يجعل تغيرات الحركة تتقدم أبطأ مقارنةً بمدة الحركة.

# --instructions--

لرؤية تأثير هذا المنحنى في العمل، غيّر خاصية `animation-timing-function` للعنصر الذي يحمل المعرف `red` إلى دالة `cubic-bezier` بقيم x1, y1, x2, y2 مضبوطة على التوالي إلى `0, 0, 0.58, 1`. هذا سيجعل كلا العنصرين يتقدمان في الحركة بشكل مشابه.

# --hints--

يجب أن تكون قيمة خاصية `animation-timing-function` للعنصر الذي يحمل المعرف `red` دالة `cubic-bezier` بقيم x1, y1, x2, y2 مضبوطة على التوالي إلى 0, 0, 0.58, 1.

```js
const redElement = document.querySelector('#red');
const redStyle = window.getComputedStyle(redElement);
assert.equal(
  redStyle?.animationTimingFunction, 'cubic-bezier(0, 0, 0.58, 1)'
);
```

يجب ألا يحتوي العنصر الذي يحمل المعرف `red` على خاصية `animation-timing-function` بقيمة `linear` بعد الآن.

```js
const redElement = document.querySelector('#red');
const redStyle = window.getComputedStyle(redElement);
assert.notEqual(redStyle?.animationTimingFunction, 'linear');
```

يجب ألا تتغير قيمة خاصية `animation-timing-function` للعنصر الذي يحمل المعرف `blue`.

```js
const blueElement = document.querySelector('#blue');
const blueStyle = window.getComputedStyle( blueElement);
const blueBallAnimation = __helpers.removeWhiteSpace(
  blueStyle?.animationTimingFunction
);
assert.isTrue(
  blueBallAnimation == 'ease-out' ||
    blueBallAnimation == 'cubic-bezier(0,0,0.58,1)'
);
```

# --seed--

## --seed-contents--

```html
<style>
  .balls{
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 27%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 56%;
    animation-timing-function: ease-out;
  }
  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id= "red"></div>
<div class="balls" id= "blue"></div>
```

# --solutions--

```html
<style>
  .balls{
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 27%;
    animation-timing-function: cubic-bezier(0, 0, 0.58, 1);
  }
  #blue {
    background: blue;
    left: 56%;
    animation-timing-function: ease-out;
  }
  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id= "red"></div>
<div class="balls" id= "blue"></div>
```
