---
id: 587d781e367417b2b2512acc
title: تثبيت عنصر في نافذة المتصفح باستخدام التمركز الثابت
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MDNUR'
forumTopicId: 301061
dashedName: lock-an-element-to-the-browser-window-with-fixed-positioning
---

# --description--

مخطط التخطيط التالي الذي يقدمه CSS هو الوضع `fixed`، وهو نوع من التمركز المطلق الذي يثبت العنصر بالنسبة لنافذة المتصفح. يشبه التمركز المطلق، يُستخدم مع خصائص إزاحة CSS كما يزيل العنصر من التدفق الطبيعي للوثيقة. العناصر الأخرى لم تعد "تدرك" مكان وجوده، مما قد يتطلب بعض التعديلات في التخطيط في أماكن أخرى.

الفرق الأساسي بين وضعي `fixed` و`absolute` هو أن العنصر ذو الوضع الثابت لا يتحرك عند تمرير المستخدم الصفحة.

# --instructions--

شريط التنقل في الشفرة مُعلَّم بمعرف `navbar`. غيّر `position` إلى `fixed`، وضع إزاحة 0 بكسل من `top` و0 بكسل من `left`. بعد إضافة الشفرة، مرّر نافذة المعاينة لترى كيف يبقى شريط التنقل في مكانه.

# --hints--

يجب أن يكون للعنصر `#navbar` خاصية `position` مضبوطة على `fixed`.

```js
const navbarElement = document.querySelector('#navbar');
const navbarStyle = window.getComputedStyle(navbarElement);
 assert.equal(navbarStyle?.position, 'fixed');
```

يجب أن تستخدم شفرتك إزاحة CSS `top` بقيمة 0 بكسل على العنصر `#navbar`.

```js
const navbarElement = document.querySelector('#navbar');
 const navbarStyle = window.getComputedStyle(navbarElement);
assert.equal(navbarStyle?.top, '0px');
```

يجب أن تستخدم شفرتك إزاحة CSS `left` بقيمة 0 بكسل على العنصر `#navbar`.

```js
const navbarElement = document.querySelector('#navbar');
 const navbarStyle = window.getComputedStyle(navbarElement);
 assert.equal(navbarStyle?.left, '0px');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    min-height: 150vh;
  }
  #navbar {



    width: 100%;
    background-color: #767676;
  }
  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }
  nav li {
    display: inline;
    margin-right: 20px;
  }
  a {
    text-decoration: none;
  }
</style>
<body>
  <header>
    <h1>Welcome!</h1>
    <nav id="navbar">
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Contact</a></li>
      </ul>
    </nav>
  </header>
  <p>I shift up when the #navbar is fixed to the browser window.</p>
</body>
```

# --solutions--

```html
<style>
  body {
    min-height: 150vh;
  }
  #navbar {
    position: fixed;
    top: 0;
    left: 0;
    width: 100%;
    background-color: #767676;
  }
  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }
  nav li {
    display: inline;
    margin-right: 20px;
  }
  a {
    text-decoration: none;
  }
</style>
<body>
  <header>
    <h1>Welcome!</h1>
    <nav id="navbar">
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Contact</a></li>
      </ul>
    </nav>
  </header>
  <p>I shift up when the #navbar is fixed to the browser window.</p>
</body>
```
