---
id: 587d781e367417b2b2512aca
title: تحريك عنصر محدد الموضع نسبيًا باستخدام إزاحات CSS
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bQEA4'
forumTopicId: 301065
dashedName: move-a-relatively-positioned-element-with-css-offsets
---

# --description--

تُخبر إزاحات CSS مثل `top` أو `bottom`، و`left` أو `right` المتصفح بمدى إزاحة عنصر ما بالنسبة إلى مكانه في التدفق الطبيعي للوثيقة. أنت تُزيح العنصر بعيدًا عن موضع معين، مما يحرك العنصر بعيدًا عن الجانب المرجعي (أي في الاتجاه المعاكس). كما رأيت في التحدٍّ السابق، استخدام إزاحة `top` حرك العنصر `h2` إلى الأسفل. وبالمثل، استخدام إزاحة `left` يحرك العنصر إلى اليمين.

# --instructions--

استخدم إزاحات CSS لتحريك العنصر `h2` 15 بكسل إلى اليمين و10 بكسل إلى الأعلى.

# --hints--

يجب أن يستخدم شفرتك إزاحة CSS لتحديد موضع العنصر `h2` نسبيًا 10 بكسل للأعلى. بعبارة أخرى، حركه 10 بكسل بعيدًا عن `bottom` من مكانه الطبيعي.

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.bottom, '10px');
```

يجب أن يستخدم شفرتك إزاحة CSS لتحديد موضع العنصر `h2` نسبيًا 15 بكسل نحو اليمين. بعبارة أخرى، حركه 15 بكسل بعيدًا عن `left` من مكانه الطبيعي.

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.left , '15px');
```

# --seed--

## --seed-contents--

```html
<head>
<style>
  h2 {
    position: relative;


  }
</style>
</head>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```

# --solutions--

```html
<head>
<style>
  h2 {
    position: relative;
    left: 15px;
    bottom: 10px;
  }
</style>
</head>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```
