---
id: 587d78a8367417b2b2512ae7
title: تغيير توقيت الحركة باستخدام الكلمات المفتاحية
challengeType: 0
videoUrl: 'https://scrimba.com/c/cJKvwCM'
forumTopicId: 301045
dashedName: change-animation-timing-with-keywords
---

# --description--

في تحريكات CSS، تتحكم خاصية `animation-timing-function` في سرعة تغير العنصر المتحرك خلال مدة التحريك. إذا كان التحريك عبارة عن سيارة تتحرك من النقطة A إلى النقطة B في وقت محدد (وهو `animation-duration` الخاص بك)، فإن `animation-timing-function` تحدد كيف تتسارع السيارة وتتباطأ خلال الرحلة.

هناك عدد من الكلمات المفتاحية المعرفة مسبقًا للاختيارات الشائعة. على سبيل المثال، القيمة الافتراضية هي `ease`، التي تبدأ ببطء، ثم تسرع في المنتصف، ثم تبطئ مجددًا في النهاية. تشمل الخيارات الأخرى `ease-out`، التي تكون سريعة في البداية ثم تبطئ، و`ease-in`، التي تبدأ ببطء ثم تسرع في النهاية، أو `linear`، التي تطبق سرعة حركة ثابتة طوال الوقت.

# --instructions--

للعناصر التي تحمل المعرف `ball1` و`ball2`، أضف خاصية `animation-timing-function` لكل منهما، واضبط `#ball1` على `linear`، و`#ball2` على `ease-out`. لاحظ الفرق في حركة العناصر أثناء التحريك لكنها تنتهي معًا، لأنهما يشتركان في نفس `animation-duration` وهو ثانيتان.

# --hints--

يجب أن تكون قيمة خاصية `animation-timing-function` للعنصر الذي يحمل المعرف `ball1` هي `linear`.

```js
const ballOne =document.querySelector('#ball1'); 
const ballOneStyle = window.getComputedStyle(ballOne); 

const ball1Animation = __helpers.removeWhiteSpace(ballOneStyle?.animationTimingFunction);
assert.isTrue(ball1Animation == 'linear' || ball1Animation == 'cubic-bezier(0,0,1,1)');
```

يجب أن تكون قيمة خاصية `animation-timing-function` للعنصر الذي يحمل المعرف `ball2` هي `ease-out`.

```js
const ballTwo = document.querySelector('#ball2'); 
const ballTwoStyle = window.getComputedStyle(ballTwo); 

const ball2Animation = __helpers.removeWhiteSpace(ballTwoStyle?.animationTimingFunction);
assert.isTrue(ball2Animation == 'ease-out' || ball2Animation == 'cubic-bezier(0,0,0.58,1)');
```

# --seed--

## --seed-contents--

```html
<style>

  .balls {
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left:27%;

  }
  #ball2 {
    left:56%;

  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }

</style>

<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```

# --solutions--

```html
<style>
  .balls {
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left:27%;
    animation-timing-function: linear;
  }
  #ball2 {
    left:56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```
