---
id: 587d78a3367417b2b2512ad1
title: تعرّف على الألوان التكاملية
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MD3Tr'
forumTopicId: 301056
dashedName: learn-about-complementary-colors
---

# --description--

نظرية اللون وتأثيرها على التصميم موضوع عميق، وتُغطى الأساسيات فقط في التحديات التالية. على موقع الويب، يمكن للون أن يجذب الانتباه إلى المحتوى، ويثير المشاعر، أو يخلق تناغمًا بصريًا. استخدام تركيبات مختلفة من الألوان يمكن أن يغير مظهر الموقع بشكل كبير، ويُفكّر كثيرًا في اختيار لوحة ألوان تناسب المحتوى.

عجلة الألوان أداة مفيدة لتصوّر كيف ترتبط الألوان ببعضها - هي دائرة حيث الألوان المتشابهة تكون متجاورة والألوان المختلفة أبعد عن بعضها. عندما يكون لونان متقابلين على العجلة، يُطلق عليهما ألوان تكاملية. تتميز هذه الألوان بأنه إذا تم دمجها، فإنها "تلغي" بعضها وتنتج لونًا رماديًا. لكن عند وضعها جنبًا إلى جنب، تظهر هذه الألوان أكثر حيوية وتنتج تباينًا بصريًا قويًا.

بعض الأمثلة على الألوان التكاملية مع رموزها الست عشرية هي:

<blockquote>الأحمر (#FF0000) والسماوي (#00FFFF)<br>الأخضر (#00FF00) والماجنتا (#FF00FF)<br>الأزرق (#0000FF) والأصفر (#FFFF00)</blockquote>

هذا يختلف عن نموذج الألوان القديم RYB الذي تعلمناه في المدرسة، والذي يحتوي على ألوان أولية وتكاملية مختلفة. تستخدم نظرية اللون الحديثة نموذج RGB الإضافي (كما على شاشة الكمبيوتر) ونموذج CMY(K) الطرحي (كما في الطباعة).

هناك العديد من أدوات اختيار الألوان المتاحة على الإنترنت التي تحتوي على خيار لإيجاد اللون التكميلي.

**ملاحظة:** استخدام اللون يمكن أن يكون وسيلة قوية لإضافة اهتمام بصري إلى الصفحة. مع ذلك، لا يجب استخدام اللون وحده كطريقة وحيدة لنقل معلومات مهمة لأن المستخدمين ذوي الإعاقة البصرية قد لا يفهمون ذلك المحتوى. ستُغطى هذه المسألة بتفصيل أكبر في تحديات إمكانية الوصول التطبيقية.

# --instructions--

غيّر خاصية `background-color` للصنفين `blue` و `yellow` إلى ألوانهما المناسبة. لاحظ كيف تبدو الألوان مختلفة بجانب بعضها مقارنةً بخلفية بيضاء.

# --hints--

عنصر `div` الذي يحمل الصنف `blue` يجب أن تكون خاصية `background-color` له باللون الأزرق.

```js
const blueElement = document.querySelector('.blue');
const blueStyle = window.getComputedStyle(blueElement); 
assert.equal(blueStyle?.backgroundColor, 'rgb(0, 0, 255)');
```

عنصر `div` الذي يحمل الصنف `yellow` يجب أن تكون خاصية `background-color` له باللون الأصفر.

```js
const yellowElement = document.querySelector('.yellow');
const yellowStyle = window.getComputedStyle(yellowElement);
assert.equal(yellowStyle?.backgroundColor, 'rgb(255, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }
  .blue {
    background-color: #000000;
  }
  .yellow {
    background-color: #000000;
  }
  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>
<div class="blue"></div>
<div class="yellow"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }
  .blue {
    background-color: blue;
  }
  .yellow {
    background-color: yellow;
  }
  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>
<div class="blue"></div>
<div class="yellow"></div>
```
