---
id: 587d78a3367417b2b2512acf
title: تغيير موضع العناصر المتداخلة باستخدام خاصية z-index
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM94aHk'
forumTopicId: 301046
dashedName: change-the-position-of-overlapping-elements-with-the-z-index-property
---

# --description--

عندما تُوضع العناصر لتتداخل (أي باستخدام `position: absolute | relative | fixed | sticky`)، يظهر العنصر الذي يأتي لاحقًا في ترميز HTML افتراضيًا فوق العناصر الأخرى. مع ذلك، تحدد خاصية `z-index` ترتيب تكديس العناصر فوق بعضها. يجب أن تكون قيمة هذه الخاصية عددًا صحيحًا (أي عددًا كاملاً وليس عشريًا)، والقيم الأعلى لـ `z-index` ترفع العنصر في التكديس فوق العناصر ذات القيم الأقل.

# --instructions--

أضف خاصية `z-index` إلى العنصر الذي يحمل اسم الصنف `first` (المستطيل الأحمر) واجعل قيمتها 2 ليغطي العنصر الآخر (المستطيل الأزرق).

# --hints--

يجب أن يكون للعنصر الذي يحمل الصنف `first` قيمة `z-index` تساوي 2.

```js
const firstElement = document.querySelector('.first');
const firstStyle = window.getComputedStyle(firstElement); 
assert.equal(firstStyle?.zIndex, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 60%;
    height: 200px;
    margin-top: 20px;
  }

  .first {
    background-color: red;
    position: absolute;

  }
  .second {
    background-color: blue;
    position: absolute;
    left: 40px;
    top: 50px;
    z-index: 1;
  }
</style>

<div class="first"></div>
<div class="second"></div>
```

# --solutions--

```html
<style>
  div {
    width: 60%;
    height: 200px;
    margin-top: 20px;
  }

  .first {
    background-color: red;
    position: absolute;
    z-index: 2;
  }
  .second {
    background-color: blue;
    position: absolute;
    left: 40px;
    top: 50px;
    z-index: 1;
  }
</style>
<div class="first"></div>
<div class="second"></div>
```
