---
id: 587d781c367417b2b2512ac2
title: ضبط حجم الخط لعناصر العناوين المتعددة
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPpQNT3'
forumTopicId: 301067
dashedName: set-the-font-size-for-multiple-heading-elements
---

# --description--

تُستخدم خاصية `font-size` لتحديد حجم النص داخل عنصر معين. يمكن تطبيق هذه القاعدة على عدة عناصر لإنشاء تناسق بصري للنص على الصفحة. في هذا التحدٍّ، ستضبط القيم لجميع الوسوم من `h1` إلى `h6` لموازنة أحجام العناوين.

# --instructions--

في وسوم `style`، اضبط `font-size` لـ:

- وسم `h1` إلى 68px.
- وسم `h2` إلى 52px.
- وسم `h3` إلى 40px.
- وسم `h4` إلى 32px.
- وسم `h5` إلى 21px.
- وسم `h6` إلى 14px.

# --hints--

يجب أن تضبط شفرتك خاصية `font-size` للوسم `h1` إلى 68 بكسل.

```js
 const fontSizeOfh1 = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh1 ,'68px');
```

يجب أن تضبط شفرتك خاصية `font-size` للوسم `h2` إلى 52 بكسل.

```js
 const fontSizeOfh2 = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh2 ,'52px');
```

يجب أن تضبط شفرتك خاصية `font-size` للوسم `h3` إلى 40 بكسل.

```js
 const fontSizeOfh3 = new __helpers.CSSHelp(document).getStyle('h3')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh3 ,'40px');
```

يجب أن تضبط شفرتك خاصية `font-size` للوسم `h4` إلى 32 بكسل.

```js
 const fontSizeOfh4 = new __helpers.CSSHelp(document).getStyle('h4')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh4 , '32px');
```

يجب أن تضبط شفرتك خاصية `font-size` للوسم `h5` إلى 21 بكسل.

```js
 const fontSizeOfh5 = new __helpers.CSSHelp(document).getStyle('h5')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh5 ,'21px');
```

يجب أن تضبط شفرتك خاصية `font-size` للوسم `h6` إلى 14 بكسل.

```js
 const fontSizeOfh6 = new __helpers.CSSHelp(document).getStyle('h6')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh6 , '14px');
```

# --seed--

## --seed-contents--

```html
<style>






</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```

# --solutions--

```html
<style>
  h1 {
    font-size: 68px;
  }
  h2 {
    font-size: 52px;
  }
  h3 {
    font-size: 40px;
  }
  h4 {
    font-size: 32px;
  }
  h5 {
    font-size: 21px;
  }
  h6 {
    font-size: 14px;
  }
</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```
