---
id: 587d78a5367417b2b2512ada
title: استخدام خاصية CSS Transform scale لتكبير عنصر عند التمرير فوقه
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLPJuM'
forumTopicId: 301077
dashedName: use-the-css-transform-scale-property-to-scale-an-element-on-hover
---

# --description--

تحتوي خاصية `transform` على وظائف متعددة تتيح لك تكبير، تحريك، تدوير، إمالة، وغيرها من التعديلات على العناصر. عند استخدامها مع الفئات المزيفة مثل `:hover` التي تحدد حالة معينة للعنصر، يمكن لخاصية `transform` إضافة تفاعلية بسهولة لعناصر الصفحة.

إليك مثالًا لتكبير عناصر الفقرة إلى 2.1 ضعف حجمها الأصلي عندما يمرر المستخدم مؤشر الفأرة فوقها:

```css
p:hover {
  transform: scale(2.1);
}
```

**ملاحظة:** عند تطبيق تحويل على عنصر `div`، سيؤثر ذلك أيضًا على جميع العناصر الفرعية الموجودة داخل هذا العنصر.

# --instructions--

أضف قاعدة CSS لحالة `hover` لعنصر `div` واستخدم خاصية `transform` لتكبير عنصر الـ `div` إلى 1.1 ضعف حجمه الأصلي عندما يمرر المستخدم مؤشر الفأرة فوقه.

# --hints--

يجب أن يتكبير حجم عنصر الـ `div` بمقدار 1.1 مرة عندما يمرر المستخدم مؤشر الفأرة فوقه.

```js
assert.match(code, /div:hover\s*?{\s*?transform:\s*?scale\(1\.1\);/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
    background: linear-gradient(
      53deg,
      #ccfffc,
      #ffcccf
    );
  }



</style>

<div></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
    background: linear-gradient(
      53deg,
      #ccfffc,
      #ffcccf
    );
  }
  div:hover {
    transform: scale(1.1);
  }
</style>
<div></div>
```
