---
id: 587d78a6367417b2b2512adb
title: استخدم خاصية CSS transform مع skewX لإمالة عنصر على المحور X
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLP8Sr'
forumTopicId: 301074
dashedName: use-the-css-transform-property-skewx-to-skew-an-element-along-the-x-axis
---

# --description--

الدالة التالية في خاصية `transform` هي `skewX()`، التي تميل العنصر المحدد على محوره السيني (الأفقي) بدرجة معينة.

الشفرة التالية تميل عنصر الفقرة بمقدار -32 درجة على المحور السيني.

```css
p {
  transform: skewX(-32deg);
}
```

# --instructions--

مِل العنصر الذي يحمل المعرف `bottom` بمقدار 24 درجة على المحور السيني باستخدام خاصية `transform`.

# --hints--

يجب أن يُميل العنصر الذي يحمل المعرف `bottom` بمقدار 24 درجة على محوره السيني.

```js
assert.match(code, /#bottom\s*?{\s*?.*?\s*?transform:\s*?skewX\(24deg\);/g);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
  }
  #top {
    background-color: red;
  }
  #bottom {
    background-color: blue;

  }
</style>

<div id="top"></div>
<div id="bottom"></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
  }
  #top {
    background-color: red;
  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>
<div id="top"></div>
<div id="bottom"></div>
```
