---
id: 587d78a7367417b2b2512ae2
title: إنشاء اتجاه بصري بتلاشي عنصر من اليسار إلى اليمين
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJqqAE'
forumTopicId: 301054
dashedName: create-visual-direction-by-fading-an-element-from-left-to-right
---

# --description--

في هذا التحدٍّ، ستغيّر `opacity` لعنصر متحرك ليختفي تدريجيًا عند وصوله إلى الجانب الأيمن من الشاشة.

في الرسوم المتحركة المعروضة، يتحرك العنصر الدائري ذو الخلفية المتدرجة إلى اليمين حتى علامة 50% من مدة الحركة وفقًا لقاعدة `@keyframes`.

# --instructions--

استهدف العنصر الذي يحمل المعرف `ball` وأضف خاصية `opacity` بقيمة 0.1 عند `50%`، ليختفي العنصر تدريجيًا أثناء تحركه إلى اليمين.

# --hints--

يجب أن تحدد قاعدة `keyframes` الخاصة بالتلاشي خاصية `opacity` بقيمة 0.1 عند 50%.

```js
assert.match(code,/@keyframes fade\s*?{\s*?50%\s*?{\s*?(?:left:\s*?60%;\s*?opacity:\s*?0?\.1;|opacity:\s*?0?\.1;\s*?left:\s*?60%;)/gi
);
```

# --seed--

## --seed-contents--

```html
<style>

  #ball {
    width: 70px;
    height: 70px;
    margin: 50px auto;
    position: fixed;
    left: 20%;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: fade;
    animation-duration: 3s;
  }

  @keyframes fade {
    50% {
      left: 60%;

    }
  }

</style>

<div id="ball"></div>
```

# --solutions--

```html
<style>
  #ball {
    width: 70px;
    height: 70px;
    margin: 50px auto;
    position: fixed;
    left: 20%;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: fade;
    animation-duration: 3s;
  }

  @keyframes fade {
    50% {
      left: 60%;
      opacity: 0.1;
    }
  }
</style>
<div id="ball"></div>
```
