---
id: 587d781b367417b2b2512abc
title: ضبط خاصية background-color للنص
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDqwA6'
forumTopicId: 301032
dashedName: adjust-the-background-color-property-of-text
---

# --description--

بدلًا من تعديل الخلفية العامة أو لون النص لجعل المقدمة سهلة القراءة، يمكنك إضافة `background-color` إلى العنصر الذي يحتوي النص الذي تريد تمييزه. يستخدم هذا التحدٍّ `rgba()` بدلًا من رموز `hex` أو `rgb()` العادية.

<blockquote>تعني rgba:<br>  r = الأحمر<br>  g = الأخضر<br>  b = الأزرق<br>  a = ألفا/مستوى الشفافية</blockquote>

يمكن أن تتراوح قيم RGB بين 0 و255. ويمكن أن تتراوح قيمة الألفا بين 1، وهي تعني لونًا صلبًا أو غير شفاف تمامًا، و0، وهي تعني شفافًا تمامًا أو مسحًا. `rgba()` ممتازة للاستخدام هنا لأنها تتيح لك تعديل الشفافية. هذا يعني أنك لا تحتاج إلى حجب الخلفية بالكامل.

ستستخدم `background-color: rgba(45, 45, 45, 0.1)` في هذا التحدٍّ. ينتج لونًا رماديًا داكنًا يكاد يكون شفافًا بسبب قيمة الشفافية المنخفضة 0.1.

# --instructions--

لجعل النص يبرز أكثر، عدل خاصية `background-color` لعنصر `h4` إلى القيمة المعطاة في `rgba()`.

أيضًا بالنسبة إلى `h4`، احذف خاصية `height` وأضف حشو `padding` بقيمة 10px.

# --hints--

يجب أن تضيف شفرتك خاصية `background-color` لعنصر `h4` وتضبطها على `rgba(45, 45, 45, 0.1)`.

```js
assert(
  /(background-color|background):rgba\(45,45,45,0?\.1\)(;?}|;)/gi.test(
    code.replace(/\s/g, '')
  )
);
```

يجب أن تضيف شفرتك خاصية `padding` لعنصر `h4` وتضبطها على 10 بكسل.

```js
const h4Element = document.querySelector("h4"); 
const h4style = window.getComputedStyle(h4Element);
assert.equal(h4style?.paddingTop, "10px");
assert.equal(h4style?.paddingRight, "10px");
assert.equal(h4style?.paddingBottom, "10px");
assert.equal(h4style?.paddingLeft, "10px");
```

يجب إزالة خاصية `height` من عنصر `h4`.

```js
const h4Element = document.querySelector("h4"); 
const h4style = window.getComputedStyle(h4Element);
assert.notEqual(h4style?.height, '25px');
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;


  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Alphabet</h4>
      <hr>
      <p><em>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
    padding: 10px;
    background-color: rgba(45, 45, 45, 0.1);
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Alphabet</h4>
      <hr>
      <p><em>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
