---
id: 587d78a3367417b2b2512ace
title: دفع العناصر إلى اليسار أو اليمين باستخدام خاصية float
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MDqu2'
forumTopicId: 301066
dashedName: push-elements-left-or-right-with-the-float-property
---

# --description--

أداة التمركز التالية لا تستخدم `position` فعليًا، بل تضبط خاصية `float` لعنصر ما. تُزال العناصر العائمة من التدفق الطبيعي للوثيقة وتُدفع إلى إما `left` أو `right` داخل العنصر الحاوي لها. تُستخدم عادةً مع خاصية `width` لتحديد مقدار المساحة الأفقية التي يحتاجها العنصر العائم.

# --instructions--

الترميز المعطى يعمل جيدًا كتصميم عمودين، مع وجود عنصري `section` و `aside` بجانب بعضهما. أعطِ العنصر الذي له المعرف `#left` خاصية `float` بقيمة `left`، والعنصر الذي له المعرف `#right` خاصية `float` بقيمة `right`.

# --hints--

يجب أن يكون للعنصر الذي يحمل المعرف `left` قيمة `float` تساوي `left`.

```js
const leftElement = document.querySelector('#left');
const leftStyle = window.getComputedStyle(leftElement);
assert.equal(leftStyle?.float, 'left');
```

يجب أن يكون للعنصر الذي يحمل المعرف `right` قيمة `float` تساوي `right`.

```js
const rightElement = document.querySelector('#right');
const rightStyle = window.getComputedStyle(rightElement);
assert.equal(rightStyle?.float, 'right');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
    #left {

      width: 50%;
    }
    #right {

      width: 40%;
    }
    aside, section {
      padding: 2px;
      background-color: #ccc;
    }
  </style>
</head>
<body>
  <header>
    <h1>Welcome!</h1>
  </header>
  <section id="left">
    <h2>Content</h2>
    <p>Good stuff</p>
  </section>
  <aside id="right">
    <h2>Sidebar</h2>
    <p>Links</p>
  </aside>
</body>
```

# --solutions--

```html
<head>
  <style>
    #left {
      float: left;
      width: 50%;
    }
    #right {
      float: right;
      width: 40%;
    }
    aside, section {
      padding: 2px;
      background-color: #ccc;
    }
  </style>
</head>
<body>
  <header>
    <h1>Welcome!</h1>
  </header>
  <section id="left">
    <h2>Content</h2>
    <p>Good stuff</p>
  </section>
  <aside id="right">
    <h2>Sidebar</h2>
    <p>Links</p>
  </aside>
</body>
```
