---
id: 587d78a4367417b2b2512ad4
title: ضبط درجة لون اللون
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp38TZ'
forumTopicId: 301036
dashedName: adjust-the-hue-of-a-color
---

# --description--

للألوان عدة خصائص تشمل الدرجة، التشبع، وفاتح اللون. قدم CSS3 الدالة `hsl()` كطريقة بديلة لاختيار لون عبر تحديد هذه الخصائص مباشرة.

**الدرجة** هي ما يفكر فيه الناس عادةً على أنه "لون". إذا تخيلت طيفًا من الألوان يبدأ بالأحمر على اليسار، ويمر بالأخضر في الوسط، وينتهي بالأزرق على اليمين، فالدرجة هي موقع اللون على هذا الخط. في `hsl()`، تستخدم الدرجة مفهوم عجلة الألوان بدل الطيف، حيث تُعطى زاوية اللون على الدائرة كقيمة بين 0 و360.

**التشبع** هو مقدار الرمادي في اللون. اللون المشبع بالكامل لا يحتوي على رمادي، واللون قليل التشبع يكاد يكون رماديًا بالكامل. يُعطى هذا كنسبة مئوية حيث 100% تعني تشبع كامل.

**فاتح اللون** هو مقدار الأبيض أو الأسود في اللون. تُعطى نسبة مئوية تتراوح بين 0% (أسود) و100% (أبيض)، حيث 50% هو اللون الطبيعي.

إليك بعض الأمثلة على استخدام `hsl()` مع ألوان مشبعة بالكامل وفاتح لون طبيعي:

<table><thead><tr><th>اللون</th><th>HSL</th></tr></thead><tbody><tr><td>red</td><td>hsl(0, 100%, 50%)</td></tr><tr><td>yellow</td><td>hsl(60, 100%, 50%)</td></tr><tr><td>green</td><td>hsl(120, 100%, 50%)</td></tr><tr><td>cyan</td><td>hsl(180, 100%, 50%)</td></tr><tr><td>blue</td><td>hsl(240, 100%, 50%)</td></tr><tr><td>magenta</td><td>hsl(300, 100%, 50%)</td></tr></tbody></table>

# --instructions--

غيّر `background-color` لكل عنصر `div` بناءً على أسماء الأصناف (`green`، `cyan`، أو `blue`) باستخدام `hsl()`. يجب أن تكون الثلاثة مشبعة بالكامل وفاتح اللون طبيعي.

# --hints--

يجب أن يستخدم شفرتك الدالة `hsl()` لإعلان اللون الأخضر.

```js
assert.match(code,/\.green\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

يجب أن يستخدم شفرتك الدالة `hsl()` لإعلان اللون السماوي.

```js
assert.match(code,/\.cyan\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

يجب أن يستخدم شفرتك الدالة `hsl()` لإعلان اللون الأزرق.

```js
assert.match(code,/\.blue\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

يجب أن يحتوي عنصر `div` الذي يحمل الصنف `green` على `background-color` باللون الأخضر.

```js
const greenElement = document.querySelector(".green");
const greenStyle = window.getComputedStyle(greenElement); 
assert.equal(greenStyle?.backgroundColor, 'rgb(0, 255, 0)');
```

يجب أن يحتوي عنصر `div` الذي يحمل الصنف `cyan` على `background-color` بلون cyan.

```js
const cyanElement = document.querySelector(".cyan");
const cyanStyle = window.getComputedStyle(cyanElement); 
assert.equal(cyanStyle?.backgroundColor, 'rgb(0, 255, 255)');
```

يجب أن يكون لعنصر `div` الذي يحمل الصنف `blue` خاصية `background-color` باللون الأزرق.

```js
const blueElement = document.querySelector(".blue");
const blueStyle = window.getComputedStyle(blueElement); 
assert.equal(blueStyle?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .green {
    background-color: #000000;
  }

  .cyan {
    background-color: #000000;
  }

  .blue {
    background-color: #000000;
  }

  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>

<div class="green"></div>
<div class="cyan"></div>
<div class="blue"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .green {
    background-color: hsl(120, 100%, 50%);
  }

  .cyan {
    background-color:   hsl(180, 100%, 50%);
  }

  .blue {
    background-color: hsl(240, 100%, 50%);
  }

  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>
<div class="green"></div>
<div class="cyan"></div>
<div class="blue"></div>
```
