---
id: 587d781d367417b2b2512ac8
title: ضبط حالة التحويم لوسم الرابط أو وسم الرابط للنص العالي
challengeType: 0
videoUrl: 'https://scrimba.com/c/cakRGcm'
forumTopicId: 301035
dashedName: adjust-the-hover-state-of-an-anchor-tag
---

# --description--

سيتناول هذا التحدٍّ استخدام الفئات المزيفة. الفئة المزيفة هي كلمة مفتاحية يمكن إضافتها إلى المحددات لاختيار حالة معينة من العنصر.

على سبيل المثال، يمكن تغيير تنسيق وسم الرابط أو وسم الرابط للنص العالي لحالته عند التحويم باستخدام محدد الفئة المزيفة `:hover`. إليك شفرة CSS لتغيير `color` لوسم الرابط إلى الأحمر أثناء حالة التحويم:

```css
a:hover {
  color: red;
}
```

# --instructions--

يحتوي محرر الشفرة على قاعدة CSS لتنسيق جميع وسم `a` باللون الأسود. أضف قاعدة بحيث عندما يحوم المستخدم فوق وسم `a`، يصبح `color` أزرق.

# --hints--

يجب أن يبقى `color` لوسم الرابط أو وسم الرابط للنص العالي أسودًا، أضف قواعد CSS فقط لحالة `:hover`.

```js
const anchorElement = document.querySelector("a"); 
const anchorStyle = window.getComputedStyle(anchorElement);
assert.equal(anchorStyle?.color, 'rgb(0, 0, 0)');
```

يجب أن يكون لوسم الرابط `color` أزرق عند التمرير عليه.

```js
assert.match(code,
    /a:hover\s*?{\s*?color:\s*?(blue|rgba\(\s*?0\s*?,\s*?0\s*?,\s*?255\s*?,\s*?1\s*?\)|#00F|rgb\(\s*?0\s*?,\s*?0\s*?,\s*?255\s*?\))\s*?;\s*?}/gi
  );
```

# --seed--

## --seed-contents--

```html
<style>
  a {
    color: #000;
  }



</style>
<a href="https://freecatphotoapp.com/" target="_blank">CatPhotoApp</a>
```

# --solutions--

```html
<style>
  a {
    color: #000;
  }
  a:hover {
    color: rgba(0,0,255,1);
  }
</style>
<a href="https://freecatphotoapp.com/" target="_blank">CatPhotoApp</a>
```
