---
id: 587d78a8367417b2b2512ae4
title: إنشاء نبض قلب باستخدام CSS مع عدد تكرار حركة لا نهائي
challengeType: 0
videoUrl: 'https://scrimba.com/c/cDZpDUr'
forumTopicId: 301062
dashedName: make-a-css-heartbeat-using-an-infinite-animation-count
---

# --description--

إليك مثالًا آخر على حركة مستمرة باستخدام الخاصية `animation-iteration-count` التي تستخدم القلب الذي صممته في تحدٍّ سابق.

تتكون حركة نبض القلب التي تستمر لمدة ثانية واحدة من قطعتين متحركتين. تُحرّك عناصر `heart` (بما في ذلك القطع `:before` و`:after`) لتغيير الحجم باستخدام الخاصية `transform`، بينما تُحرّك خلفية العنصر `div` لتغيير لونها باستخدام الخاصية `background`.

# --instructions--

حافظ على نبض القلب بإضافة الخاصية `animation-iteration-count` لكل من الصنف `back` والصنف `heart` مع تعيين القيمة إلى `infinite`. لا يحتاج المحددان `heart:before` و`heart:after` إلى أي خصائص حركة.

# --hints--

يجب أن تكون قيمة الخاصية `animation-iteration-count` للصنف `heart` هي `infinite`.

```js
const heartElement = document.querySelector('.heart');
 const heartStyle = window.getComputedStyle(heartElement);
 assert.equal(heartStyle?.animationIterationCount, 'infinite');
```

يجب أن تكون قيمة الخاصية `animation-iteration-count` للصنف `back` هي `infinite`.

```js
const backElement = document.querySelector('.back');
 const backStyle = window.getComputedStyle(backElement);
 assert.equal(backStyle?.animationIterationCount, 'infinite');
```

# --seed--

## --seed-contents--

```html
<style>
  .back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: white;
    animation-name: backdiv;
    animation-duration: 1s;

  }

  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: rotate(-45deg);
    animation-name: beat;
    animation-duration: 1s;

  }
  .heart:after {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart:before {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }

  @keyframes backdiv {
    50% {
      background: #ffe6f2;
    }
  }

  @keyframes beat {
    0% {
      transform: scale(1) rotate(-45deg);
    }
    50% {
      transform: scale(0.6) rotate(-45deg);
    }
  }

</style>
<div class="back"></div>
<div class="heart"></div>
```

# --solutions--

```html
<style>
  .back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: white;
    animation-name: backdiv;
    animation-duration: 1s;
    animation-iteration-count: infinite;
  }

  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: rotate(-45deg);
    animation-name: beat;
    animation-duration: 1s;
    animation-iteration-count: infinite;
  }
  .heart:after {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart:before {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }

  @keyframes backdiv {
    50% {
      background: #ffe6f2;
    }
  }

  @keyframes beat {
    0% {
      transform: scale(1) rotate(-45deg);
    }
    50% {
      transform: scale(0.6) rotate(-45deg);
    }
  }
</style>
<div class="back"></div>
<div class="heart"></div>
```
