---
id: 58a7a6ebf9a6318348e2d5aa
title: تعديل وضع التعبئة للرسوم المتحركة
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJDmcE'
forumTopicId: 301064
dashedName: modify-fill-mode-of-an-animation
---

# --description--

هذا جيد، لكنه لا يعمل بشكل صحيح بعد. لاحظ كيف تعود الرسوم المتحركة إلى البداية بعد مرور `500ms`، مما يجعل الزر يعود إلى لونه الأصلي. تريد أن يبقى الزر مميزًا.

يمكنك فعل ذلك بتعيين خاصية `animation-fill-mode` إلى القيمة `forwards`. تحدد خاصية `animation-fill-mode` النمط الذي يُطبق على العنصر بعد انتهاء الرسوم المتحركة. يمكنك تعيينها كما يلي:

```css
animation-fill-mode: forwards;
```

# --instructions--

عيّن خاصية `animation-fill-mode` في `button:hover` إلى `forwards` ليبقى الزر مميزًا عندما يمرر المستخدم المؤشر فوقه.

# --hints--

يجب أن تحتوي `button:hover` على خاصية `animation-fill-mode` بقيمة `forwards`.

```js
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-fill-mode\s*?:\s*?forwards\s*?;[\s\S]*}/gi);
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-name\s*?:\s*?background-color\s*?;[\s\S]*}/gi);
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-duration\s*?:\s*?500ms\s*?;[\s\S]*}/gi );
```

# --seed--

## --seed-contents--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
    /* Only change code below this line */

    /* Only change code above this line */
  }
  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```

# --solutions--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
    animation-fill-mode: forwards;
  }
  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```
