---
id: 66f26c32ec6f90df01a44f60
title: 設計一個聯繫表單
challengeType: 25
dashedName: design-a-contact-form
demoType: onClick
---

# --description--

**目標：**實現以下用戶需求並通過所有測試以完成實驗。

**用戶需求：**

1. 你應該有一個類爲 `form-container` 的 `div`，用於放置表單。

2. 你在 `.form-container` 元素中應該有一個 `form` 元素。

3. 在 `form` 元素中，應包含以下元素和輸入字段：

   - 一個包含一些文本的 `h2` 元素。
   - 一個姓名文本輸入字段，`type` 設置爲 `text`，幷包含 `id`、`name`、`required` 屬性。
   - 一個電子郵件輸入字段，`type` 設置爲 `email`，幷包含 `id`、`name`、`required` 屬性。
   - 一個信息文本區域，包含 `id`、`name`、`required` 屬性。
   - 一個 `button` 元素，文本爲 `"Submit"`，`type` 屬性爲 `submit`。

4. 每個 `input` 和 `textarea` 元素都應該有其相應的 `label` 元素，將 `for` 屬性值設置爲元素的 `id` 值。

5. `.form-container` 元素應具有背景色，並且 `border-radius`、`padding` 和 `width` 都有值。

6. `label` 元素應有外邊距和字體顏色。

7. `input` 和 `textarea` 元素應具有以下值：

   - `width`。
   - `padding`。
   - 底部外邊距。

8. `button` 元素應具有：

   - 背景色
   - 非默認的字體大小。

9.  按鈕應具有改變背景顏色的懸停效果。

**注意：** 確保在 HTML 中鏈接樣式表以應用 CSS。

# --hints--

你應該有一個類爲 `form-container` 的 `div`。

```js
const el = document.querySelector('div.form-container');
assert.isNotNull(el);
```

你在 `.form-container` 元素中應該有一個 `form` 元素。

```js
const el = document.querySelector('.form-container form');
assert.isNotNull(el);
```

在 `form`中，你應該有一個包含一些文本的 `h2` 元素。

```js
const el = document.querySelector('.form-container form h2');
assert.isNotEmpty(el.innerText);
// assert.isAtLeast(document.querySelector('h1').innerText.length, 1);

```

在 `form`中，你應該有一個類型設置爲 `text` 的文本輸入字段。

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert.isNotNull(el);
```

你的 `type=text` 輸入字段應該有一個 `id`、一個 `name` 和一個 `required` 屬性。

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert.exists(el);
assert.isTrue(el.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

在 `form`中，你應該有一個電子郵件輸入字段，`type` 設置爲 `email`。

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert.isNotNull(el);
```

你的 `type=email` 輸入字段應該有一個 `id`、一個 `name` 和一個 `required` 屬性。

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert.isTrue(el?.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

在 `form`中，你應該有一個 `textarea`，包含一個 `id`、一個 `name` 和一個 `required` 屬性。

```js
const el = document.querySelector('.form-container form textarea');
assert.isNotNull(el);
assert.isNotNull(document.querySelector('.form-container form textarea[required]'));
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

在 `form`中，你應該有一個 `button` 元素，`type` 屬性值爲 `submit`。

```js
const el = document.querySelector('.form-container form button');
assert.strictEqual(el?.getAttribute('type'), 'submit');
```

你的按鈕文本應爲 `"Submit"` 。

```js
const el = document.querySelector('.form-container form button');
assert.strictEqual(document.querySelector('.form-container form button').textContent, 'Submit');
```

你應該有一個 `label` 元素與 `#name` 元素關聯，其 `for` 屬性與相應 `input` 的 `id` 值相匹配。

```js
const nameInput = document.querySelector('.form-container form input[type="text"]')?.getAttribute('id');

const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

// check if id has a value
assert.isNotEmpty(nameInput);
let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === nameInput) {
        found = true;
        break;
    }
}

assert.isTrue(found);
```

你應該有一個 `label` 元素與 `#email` 元素關聯，其 `for` 屬性與相應 `input` 的 `id` 值相匹配。

```js
const emailInput = document.querySelector('.form-container form input[type="email"]')?.getAttribute('id');
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);
// check if id has a value
assert.isNotEmpty(emailInput);
let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === emailInput) {
        found = true;
        break;
    }
}

assert.isTrue(found);
```

你應該有一個 `label` 元素與 `#message` 元素關聯，其 `for` 屬性與相應 `textarea` 的 `id` 值相匹配。

```js
const messageInput = document.querySelector('.form-container form textarea')?.getAttribute('id');
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

// check if id has a value

assert.isNotEmpty(messageInput);

assert.exists(document.querySelector(`label[for="${messageInput}"]`));
```

標籤不應爲空。

```js
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

for (let label of labels) {
    const labelText = label.textContent.trim();
    assert.isNotEmpty(labelText);
}

```

`.form-container` 元素應具有背景顏色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.backgroundColor);
```

`.form-container` 元素應具有 `border-radius`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.borderRadius);
```

`.form-container` 元素應具有內邊距。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.padding);
```

`.form-container` 元素應具有寬度。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.width);
```

標籤應有外邊距和字體顏色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('label')?.margin);
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('label')?.color);
```

你的 `input` 元素應有寬度、內邊距和底部外邊距。

```js
// Helper function to extract properties from multiple selectors
function getStyleProperties(selectors, properties) {
  const result = {};
  properties.forEach((property) => {
    result[property] = selectors.map((selector) => selector?.[property]).find(Boolean);
  });
  return result;
}

// Initializing selectors
const selectors = [
  new __helpers.CSSHelp(document).getStyle('input'),
  new __helpers.CSSHelp(document).getStyle('input, textarea'),
  new __helpers.CSSHelp(document).getStyle('textarea, input')
];

// Extract width, padding, and margin properties
const { width } = getStyleProperties(selectors, ['width']);
assert.isNotEmpty(width);

const {
  padding,
  paddingTop,
  paddingBottom,
  paddingLeft,
  paddingRight
} = getStyleProperties(selectors, ['padding', 'paddingTop', 'paddingBottom', 'paddingLeft', 'paddingRight']);

// Assert padding properties
assert.isOk(
  padding ||
  (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

const { marginBottom } = getStyleProperties(selectors, ['marginBottom']);
assert.isNotEmpty(marginBottom);

```

你的 `textarea` 元素應有寬度、內邊距和底部外邊距。

```js
// Helper function to extract properties from multiple selectors
function getStyleProperties(selectors, properties) {
  const result = {};
  properties.forEach((property) => {
    result[property] = selectors.map((selector) => selector?.[property]).find(Boolean);
  });
  return result;
}

// Initializing selectors
const selectors = [
  new __helpers.CSSHelp(document).getStyle('textarea'),
  new __helpers.CSSHelp(document).getStyle('input, textarea'),
  new __helpers.CSSHelp(document).getStyle('textarea, input')
];

// Extract width, padding, and margin properties
const { width } = getStyleProperties(selectors, ['width']);
assert.isNotEmpty(width);

const {
  padding,
  paddingTop,
  paddingBottom,
  paddingLeft,
  paddingRight
} = getStyleProperties(selectors, ['padding', 'paddingTop', 'paddingBottom', 'paddingLeft', 'paddingRight']);

// Assert padding properties
assert.isOk(
  padding ||
  (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

const { marginBottom } = getStyleProperties(selectors, ['marginBottom']);
assert.isNotEmpty(marginBottom);
```

你的 `button` 元素應有背景顏色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button')?.backgroundColor);
```

你的 `button` 元素應有字體大小。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button')?.fontSize);
```

你的 `button` 元素應有改變背景顏色的懸停效果。

```js
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
const btn_hover = gs('button:hover')?.backgroundColor;
assert.isNotEmpty(btn_hover);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Contact Form</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Styled Form Project</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="form-container">
        <form>
            <h2>Contact Us</h2>
            <label for="name">Name:</label>
            <input type="text" id="name" name="name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" required>

            <label for="message">Message:</label>
            <textarea id="message" name="message" required></textarea>

            <button type="submit">Submit</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background: linear-gradient(45deg, #83a4d4, #b6fbff);
    font-family: Arial, sans-serif;
}

.form-container {
    background-color: white;
    border-radius: 10px;
    padding: 20px;
    width: 300px;
    text-align: center;
}

form h2 {
    margin-bottom: 20px;
    color: #333;
}

label {
    display: block;
    margin: 10px 0 5px;
    color: #555;
}

input{
    width: 100%;
    padding: 10px;
    margin-bottom: 10px;
    border: 1px solid #ddd;
    border-radius: 5px;
    box-sizing: border-box;
    resize: vertical;
}

textarea {
    width: 100%;
    padding: 10px;
    margin-bottom: 10px;
    border: 1px solid #ddd;
    border-radius: 5px;
    box-sizing: border-box;
    resize: vertical;
}

button {
    background-color: #4CAF50;
    color: white;
    padding: 10px 15px;
    border: none;
    border-radius: 5px;
    cursor: pointer;
    font-size: 1em;
}

button:hover {
    background-color: #45a049;
}
```
