---
id: ab306dbdcc907c7ddfc30830
title: 創建一個深度扁平化工具
challengeType: 26
dashedName: create-a-deep-flattening-tool
---

# --description--

在這個實驗中，你將實作一個陣列扁平化演算法。

將陣列攤平成一維陣列表示將任意深度的巢狀陣列轉換成單一的一維陣列。此處理程序會依序擷取所有元素，僅展開陣列。其他型別則保持不變。

例如：

| Original | Flattened |
| --- | --- |
| `[[1], [], [2, [3]]]` | `[1, 2, 3]` |
| `[1, {"foo": "bar"}, [2]]` | `[1, {"foo": "bar"}, 2]` |
| `["baz", [1, 2], {}]` | `["baz", 1, 2, {}]` |

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個名為 `steamrollArray` 的函式（程式）。
1. `steamrollArray` 函式（程式）應該接受一個引數：一個巢狀陣列。
1. 該函式（程式）應該將巢狀陣列攤平成一維，並考慮不同的巢狀層級。
1. 你的解決方案不應該使用 `Array.prototype.flat()` 或 `Array.prototype.flatMap()` 方法。
1. 不應該使用全域變數。

# --hints--

你應該有一個 `steamrollArray` 函式（程式）。

```js
assert.isFunction(steamrollArray);
```

`steamrollArray([[["a"]], [["b"]]])` 應該傳回 `["a", "b"]`。

```js
assert.deepEqual(steamrollArray([[['a']], [['b']]]), ['a', 'b']);
```

`steamrollArray([1, [2], [3, [[4]]]])` 應該傳回 `[1, 2, 3, 4]`。

```js
assert.deepEqual(steamrollArray([1, [2], [3, [[4]]]]), [1, 2, 3, 4]);
```

`steamrollArray([1, [], [3, [[4]]]])` 應該傳回 `[1, 3, 4]`。

```js
assert.deepEqual(steamrollArray([1, [], [3, [[4]]]]), [1, 3, 4]);
```

`steamrollArray([1, {}, [3, [[4]]]])` 應該傳回 `[1, {}, 3, 4]`。

```js
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4]);
```

你的解決方案不應該使用 `Array.prototype.flat()` 或 `Array.prototype.flatMap()` 方法。

```js
const arr = [1,2,3,[1,2,3,[1,2,3]]];

const spyFlat = __helpers.spyOn(Array.prototype, 'flat');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlat.calls);
} finally {
  spyFlat.restore();
}

const spyFlatMap = __helpers.spyOn(Array.prototype, 'flatMap');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlatMap.calls);
} finally {
  spyFlatMap.restore();
}
```

你不應該使用全域變數。

```js
steamrollArray([1, {}, [3, [[4]]]])
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4])
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function steamrollArray(arr) {
  if (!Array.isArray(arr)) {
    return [arr];
  }
  const out = [];
  arr.forEach(function(e) {
    steamrollArray(e).forEach(function(v) {
      out.push(v);
    });
  });
  return out;
}
```
