---
id: 68b94f23f7b3ee14078c89c5
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستستخدم عنصر `iframe` لعرض فيديو. تم إعداد قالب HTML جاهز لك.

ابدأ بإنشاء عنصر `h1` يحتوي على النص `iframe Video Display`.

# --hints--

يجب أن يكون لديك عنصر `h1`.

```js
assert.exists(document.querySelector('h1'))
```

يجب أن يحتوي عنصر `h1` الخاص بك على النص `iframe Video Display`.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), 'iframe Video Display')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
