---
id: 68ba9c4f1688914d72876095
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

عنصر `iframe` يأخذ أيضًا خاصية `src` بقيمة تشير إلى عنوان URL أو مسار المورد الذي سيتم عرضه.

أضف خاصية `src` بقيمة `https://www.youtube.com/embed/I0_951_MPE0` إلى عنصر `iframe` الخاص بك.

في هذه المرحلة، يجب أن ترى الفيديو معروضًا على الصفحة، لكن هناك بعض الخصائص الأخرى التي تحتاج إلى إضافتها.

# --hints--

يجب أن يحتوي عنصر `iframe` الخاص بك على خاصية `src` مضبوطة على `https://www.youtube.com/embed/I0_951_MPE0`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('src'), 'https://www.youtube.com/embed/I0_951_MPE0')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
