---
id: 68ba9c4f1688914d72876094
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

في الدرس الأول عن عنصر `iframe`، تعلمت أنه عنصر مُستبدَل مثل `img`. هذا يعني أنه يمكنه أيضًا أخذ الخاصيتين `width` و`height` لتحديد ارتفاعه وعرضه.

أعطِ عنصر `iframe` الخاص بك خاصية `width` بقيمة `560` وخاصية `height` بقيمة `315`.

# --hints--

يجب أن يحتوي عنصر `iframe` الخاص بك على خاصية `width` مضبوطة على `560`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('width'), '560')
```

يجب أن يحتوي عنصر `iframe` الخاص بك على خاصية `height` مضبوطة على `315`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('height'), '315')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>

    <iframe
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
