---
id: 68ba9c4f1688914d72876096
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

واحدة من الخصائص هي `allow`. تشبه قائمة أذونات تخبر المتصفح بالميزات التي يُسمح لـ `iframe` باستخدامها.

إليك عنصر `iframe` مع الخاصية `allow`:

```html
<iframe
  allow="accelerometer autoplay clipboard-write encrypted-media gyroscope picture-in-picture web-share"
></iframe>
```

أضف الخاصية `allow` بالقيم `accelerometer`، `autoplay`، و`clipboard-write`.

يتيح `accelerometer` لـ `iframe` استخدام حساسات الحركة ليكتشف أشياء مثل ميل الجهاز ودورانه. يتيح `autoplay` بدء تشغيل الفيديو تلقائيًا، ويتيح `clipboard-write` لـ `iframe` كتابة البيانات إلى الحافظة الخاصة بالمستخدم.

# --hints--

يجب أن يحتوي عنصر `iframe` على خاصية `allow` مضبوطة على `accelerometer`.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /accelerometer/)
```

يجب أن تحتوي الخاصية `allow` في عنصر `iframe` على `autoplay` كواحدة من قيمها.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /autoplay/)
```

يجب أن تحتوي الخاصية `allow` في عنصر `iframe` على `clipboard-write` كواحدة من قيمها.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /clipboard-write/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
