---
id: 68ba9c4f1688914d72876093
title: خطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

الآن، أنشئ عنصر `iframe`. لا تضع فيه أي شيء بعد.

# --hints--

يجب أن يكون لديك عنصر `iframe`.

```js
assert.exists(document.querySelector('iframe'))
```

يجب أن يكون عنصر `iframe` فارغًا.

```js
assert.equal(document.querySelector('iframe')?.innerHTML.trim(), '')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
