---
id: 68ba9c4f1688914d72876097
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

أضف `encrypted-media` و`gyroscope` و`web-share` إلى القيم الموجودة في الخاصية `allow`.

ستسمح هذه الثلاثة باستخدام امتدادات الوسائط المشفرة لحماية الفيديو، ومنح وصول إلى حساسات الحركة والتوجيه في الجهاز، والسماح بمشاركة محتوى `iframe` عبر نوافذ المشاركة الأصلية في الجهاز.

# --hints--

يجب أن تحتوي الخاصية `allow` في عنصر `iframe` على `encrypted-media` كواحدة من قيمها.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /encrypted-media/)
```

يجب أن تحتوي الخاصية `allow` في عنصر `iframe` على `gyroscope` كواحدة من قيمها.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /gyroscope/)
```

يجب أن تحتوي الخاصية `allow` في عنصر `iframe` على `web-share` كواحدة من قيمها.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /web-share/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      allow="accelerometer autoplay clipboard-write"
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
