---
id: 68ba9c4f1688914d72876098
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

الخاصية التالية التي ستضيفها هي `referrerpolicy`. هذه القاعدة تحدد مقدار التفاصيل التي تشاركها عندما تتصل صفحتك بصفحة أخرى.

أضف الخاصية `referrerpolicy` واضبطها على `strict-origin-when-cross-origin`. هذا يشارك العنوان الكامل على نفس الموقع، واسم الموقع فقط على المواقع الأخرى، ولا شيء على المواقع غير الآمنة.

# --hints--

يجب أن يحتوي عنصر `iframe` على خاصية `referrerpolicy` مضبوطة على `strict-origin-when-cross-origin`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('referrerpolicy'), 'strict-origin-when-cross-origin')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
      allow="accelerometer autoplay clipboard-write encrypted-media gyroscope web-share"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
