---
id: 660f23b53db70af0f2620e78
title: 第 43 步
challengeType: 1
dashedName: step-43
---

# --description--

现在，你的所有数字都出现在同一行上。 这对于创建金字塔不起作用。

你需要在每一行中添加一个新行。 但是，在 JavaScript 中按回车键在引号之间插入换行符将导致解析错误。 相反，你需要使用特殊的 <dfn>转义序列</dfn> `\n`，它在记录字符串时被解释为新行。 例如：

```js
lineOne = lineOne + "\n" + lineTwo;
```

使用第二个加法运算符在 `result` 和 `row` 值之后添加一个换行符。

# --hints--

你应该使用 `\n` 转义序列。 请记住，它需要是一个字符串，因此用引号括起来。

```js
assert.match(__helpers.removeJSComments(code), /('|")\\n\1/);
```

你应该将 `row` 串联到 `result` 变量。

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row\s*\+\s*('|")\\n\1/);
```

你应该在 `row` 变量后拼接 `\n` 转义字符。

```js
assert.match(__helpers.removeJSComments(code), /row\s*\+\s*('|")\\n\1/);
```

你应该将整个拼接结果赋值给 `result` 变量。

```js
assert.strictEqual(result, "0\n1\n2\n3\n4\n5\n6\n7\n");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
