---
id: 66643c93e05093c728abdbe9
title: 步骤 55
challengeType: 1
dashedName: step-55
---

# --description--

在继续之前，请花点时间回顾一下函数的工作原理。

声明一个名为 `addTwoNumbers` 的函数。 该函数应该接受两个参数并返回这两个参数的总和。

您的函数不应使用硬编码的值。 硬编码函数的例子是这样：

```js
function sayName(firstName, lastName) {
  return "John Doe";
}

sayName("Camper", "Cat");
```

无论传递给参数 `firstName` 和 `lastName` 的参数是什么，此函数都会返回 `"John Doe"`，因此 `"John Doe"` 被视为硬编码值。

声明一个 `sum` 变量，并为其分配使用 `5` 和 `10` 作为参数调用 `addTwoNumbers` 函数的值。 将 `sum` 变量记录到控制台。

# --hints--

你应该有一个名为 `addTwoNumbers` 的函数。

```js
assert.isFunction(addTwoNumbers);
```

你的函数 `addTwoNumbers` 应该有两个参数。

```js
assert.lengthOf(addTwoNumbers, 2);
```

你的函数应该返回两个参数的和。

```js
assert.strictEqual(addTwoNumbers(5,10), 15);
```

你的函数不应返回硬编码值。 也就是说，它应该适用于任何两个数字参数。

```js
assert.strictEqual(addTwoNumbers(3, 5), 8);
```

你应该声明一个 `sum` 变量。

```js
assert.isDefined(sum);
```

你的 `sum` 变量应该具有值 `15`。

```js
assert.strictEqual(sum, 15);
```

你应该将以 `5` 和 `10` 作为参数调用 `addTwoNumbers` 函数得到的值分配给 `sum`。

```js
assert.match(code, /sum\s*=\s*addTwoNumbers\s*\(/);
```

你应该记录你的 `sum` 变量。

```js
assert.match(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--

--fcc-editable-region--

const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
