---
id: 655cd899f8de09431eabb40c
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo progetto imparerai le strutture dati costruendo l'algoritmo di merge sort.

Si tratta di un algoritmo di ordinamento che usa il principio divide et impera per ordinare collezioni di dati. In pratica, "divide" una collezione in sotto-parti più piccole, e "conquista" le sotto-parti ordinandole indipendentemente, poi unisce le sotto-parti ordinate.

Crea una funzione chiamata `merge_sort`. Questa funzione si occuperà dell'attività di ordinare una lista di numeri.

Usa la parola chiave `pass` nel corpo della funzione.

# --hints--

Dovresti dichiarare una funzione chiamata `merge_sort`. Non dimenticare di usare la parola chiave `pass` nel corpo della funzione.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(merge_sort)
`))})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
