---
id: 6564a9fe51964c229d5b7f4c
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Ti serve un parametro che indichi la raccolta di dati da ordinare. Crea un parametro chiamato `array` nella funzione `merge_sort`.

# --hints--

La tua funzione `merge_sort` dovrebbe prendere un solo parametro: `array`.

```js
({ test: () => assert(runPython(`
    import inspect
    str(inspect.signature(merge_sort)) == '(array)'    
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort():
    pass
--fcc-editable-region--
```
