---
id: 6564aee9c077774ea49c3faf
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

L'algoritmo di ordinamento merge sort esegue principalmente tre azioni:

- Dividere una sequenza di elementi non ordinata in sotto-parti
- Ordinare gli elementi nelle sotto-parti
- Unire le sotto-parti ordinate

Quanto sopra avviene in modo ricorsivo finché le sotto-parti non vengono unite nella sequenza completa ordinata. Iniziamo dividendo la sequenza.

Per prima cosa, sostituisci la parola chiave `pass` con una variabile `middle_point`. Usa l'operatore di divisione intera (`//`) per dividere la lunghezza della lista `array` per 2 e assegna il risultato alla tua nuova variabile `middle_point`. Ricorda di indentare il codice.

# --hints--

Non dovresti avere `pass` nel corpo della funzione.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.notMatch(function_body, /^\s{4}pass/m);
  }
})
```

Dovresti dichiarare una variabile `middle_point` e assegnarle `len(array) // 2`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}middle_point\s*=\s*len\(\s*array\s*\)\s*\/\/\s*2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
