---
id: 656680b0fc79f2c38a34d90e
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

Ora è il momento di ordinare e unire le liste (`left_part` e `right_part`) nella lista originale `array`.

Puoi farlo confrontando gli elementi di entrambe le liste e unendo l'elemento più piccolo nella lista principale. Farai questo confronto per tutti gli indici in `left_part` e `right_part`.

Crea tre variabili: `left_array_index`, `right_array_index` e `sorted_index` e imposta i loro valori a `0`. Queste variabili ti aiuteranno a tenere traccia di ogni indice durante il processo di ordinamento.

# --hints--

Dovresti avere una variabile chiamata `left_array_index` all'interno della tua funzione `merge_sort`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_array_index/);
  }
})
```

Il valore di `left_array_index` dovrebbe essere 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_array_index\s*=\s*0/);
  }
})
```

Dovresti avere una variabile chiamata `right_array_index` all'interno della tua funzione `merge_sort`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_array_index/);
  }
})
```

Il valore di `right_array_index` dovrebbe essere 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_array_index\s*=\s*0/);
  }
})

```

Dovresti avere una variabile chiamata `sorted_index` all'interno della tua funzione `merge_sort`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /sorted_index/);
  }
})
```

Il valore di `sorted_index` dovrebbe essere 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /sorted_index\s*=\s*0/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)

--fcc-editable-region--

--fcc-editable-region--
```
