---
id: 656639771fed09b5c0e8fe71
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Ora che hai diviso la lista `array` in due liste separate, continuerai a dividere ogni lista finché ogni elemento non sarà da solo nella propria lista. Una lista con un solo numero è sempre ordinata.

Per farlo, chiama ricorsivamente `merge_sort` all'interno della tua funzione. Per ora non passare alcun argomento.

# --hints--

Dovresti chiamare `merge_sort` alla fine del corpo della tua funzione.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
--fcc-editable-region--
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
