---
id: 656627b47bd2d2a4afbc945d
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Nel passaggio precedente, hai ottenuto il punto medio. Puoi usarlo per dividere `array` in due parti e assegnare ciascuna parte a nuove variabili.

Usa la sintassi slice per estrarre la metà sinistra di `array` e assegnarla a una variabile chiamata `left_part`.


# --hints--

Dovresti avere una variabile chiamata `left_part` nella tua funzione `merge_sort`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*(?!=)/m);
  }
})
```

Dovresti assegnare `array[:middle_point]` alla variabile `left_part`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*=\s*array\s*\[\s*:\s*middle_point\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
