---
id: 656707e11e671ed54c96f7ec
title: Passaggio 13
challengeType: 20
dashedName: step-13
---

# --description--

Quando la condizione `if` valuta a `True`, significa che l'elemento nella lista `left_part` è più piccolo dell'elemento con cui viene confrontato nella lista `right_part`.

In quel caso, puoi assegnare l'indice `left_part` all'array ordinato.

All'interno del blocco `if`, rimuovi `pass` e assegna `left_part[left_array_index]` a `array[sorted_index]`.

# --hints--

Dovresti sostituire `pass` con `array[sorted_index] = left_part[left_array_index]` nel corpo dell'istruzione `if`.

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_body().is_equivalent("array[sorted_index] = left_part[left_array_index]")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            pass

--fcc-editable-region--
```
