---
id: 656709e50ed928da35260276
title: Passaggio 15
challengeType: 20
dashedName: step-15
---

# --description--

In un passaggio precedente, hai assegnato l'elemento nella `left_part` alla lista `array` perché era più piccolo. Ma questo non sarà sempre il caso. In alcuni casi di confronto, l'elemento a destra potrebbe essere più piccolo.

Crea una clausola `else` da eseguire quando l'indice `left_part` non è minore dell'indice `right_part`.

All'interno del blocco `else`, assegna `right_part[right_array_index]` a `array[sorted_index]`.

# --hints--

Dovresti avere `array[sorted_index] = right_part[right_array_index]` nel corpo della tua clausola `else`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /else:\s*[^}]*array\s*\[\s*sorted_index\s*\]\s*=\s*right_part\s*\[\s*right_array_index\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
--fcc-editable-region--
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
--fcc-editable-region--
```
