---
id: 656710d1e0ec62253426db24
title: Passaggio 18
challengeType: 20
dashedName: step-18
---

# --description--

Il loop `while` che hai creato confronta un elemento da `left_part` con un altro in `right_part`, poi aggiunge l'elemento più piccolo alla lista principale `array`.

Continuerà questa operazione finché non ci saranno più elementi da confrontare. Ma `left_part` potrebbe ancora avere elementi mentre `right_part` non ne ha, e viceversa.

Crea un altro loop `while` per copiare gli elementi rimanenti in `left_part` nella lista `array`. Usa `left_array_index < len(left_part)` per la condizione `while`.

# --hints--

Il tuo loop `while` dovrebbe avere questa condizione: `left_array_index < len(left_part)`. Non dimenticare di usare la parola chiave `pass` nel corpo del loop.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}while\s+left_array_index\s*<\s*len\(\s*left_part\s*\):/m);
  }
})
```

Dovresti avere la parola chiave `pass` nel corpo del tuo loop `while`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\):[^}]*\bpass\b/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
--fcc-editable-region--

--fcc-editable-region--
```
