---
id: 656702f8b4cbd8cbf0a433c6
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

All'interno della tua funzione, crea un loop `while` che confronta un elemento in `left_part` con un elemento in `right_part` e unisce l'elemento più piccolo alla lista principale `array`.

Crea due condizioni per il loop: una che verifica se `left_array_index` è minore della lunghezza di `left_part` e un'altra che verifica se `right_array_index` è minore della lunghezza di `right_part`.

# --hints--

Il tuo loop `while` dovrebbe avere queste condizioni: `left_array_index < len(left_part) and right_array_index < len(right_part)`. Non dimenticare di usare la parola chiave `pass` nel corpo del loop.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):/);
  }
})
```

Usa la parola chiave `pass` nel corpo del loop

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):\s*pass/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--

--fcc-editable-region--
```
