---
id: 6626b4c58c027d86478ff5eb
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

El profesor está muy contento con el programa que has creado hasta ahora. Pero ahora quieren tener una forma fácil de comprobar si un estudiante tiene una nota aprobada. Una nota aprobada es cualquier cosa que no sea una `"F"`.

Completa la función `hasPassingGrade` que toma una puntuación del estudiante como parámetro. Tu función debe retornar `true` si el estudiante tiene una calificación aprobada y `false` si no lo tiene.

**Consejos**

- Usa la función `getGrade` para obtener la calificación del estudiante. Luego comprueba si la nota está aprobada o no.

# --hints--

Tu función `hasPassingGrade` debe devolver un valor booleano.

```js
assert.strictEqual(typeof hasPassingGrade(100), "boolean");
```

Tu función `hasPassingGrade` debe devolver `true` si la nota es `"A"`.

```js
assert.isTrue(hasPassingGrade(100));
```

Tu función `hasPassingGrade` debe devolver `false` si la nota es `"F"`.

```js
assert.isFalse(hasPassingGrade(53));
```

Tu función `hasPassingGrade` debe devolver `true` para todos las calificaciones pasadas.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

--fcc-editable-region--
function hasPassingGrade(score) {

}


console.log(hasPassingGrade(100));
console.log(hasPassingGrade(53));
console.log(hasPassingGrade(87));
--fcc-editable-region--
```
