---
id: 6626b8dcf5057f896f948440
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Ahora que el profesor tiene toda la información que necesita, quiere poder enviar un mensaje al estudiante con los resultados.

Completa la función `studentMsg` con `totalScores` y `studentScore` como parámetros. La función debe devolver una cadena que represente un mensaje al estudiante.

Si el estudiante ha aprobado el curso, la cadena debe seguir este formato:

```md
Class average: average-goes-here. Your grade: grade-goes-here. You passed the course.
```

Si el estudiante no ha aprobado en el curso, la cadena debe seguir este formato:

```md
Class average: average-goes-here. Your grade: grade-goes-here. You failed the course.
```

Reemplace `average-goes-here` con el promedio de las calificaciones totales. Reemplace `grade-goes-here` con la calificación del estudiante.

**Consejos**

- Usa la función `getAverage` para obtener el promedio de la clase.
- Usa la función `getGrade` para obtener la calificación del estudiante.
- Usa la concatenación de cadenas (`+`) para construir el mensaje.
- Tenga cuidado con la puntuación y los espacios del mensaje.

# --hints--

Tu llamada a la función `studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` debe devolver el siguiente mensaje: `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

Tu llamada a la función `studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` debe devolver el siguiente mensaje: `"Class average: 50.8. Your grade: A++. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A++. You passed the course.");
```

Tu función `studentMsg` debe devolver el mensaje correcto basado en la puntuación del estudiante y el promedio de la clase.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

--fcc-editable-region--
function studentMsg(totalScores, studentScore) {

}
console.log(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37));
--fcc-editable-region--

```


# --solutions--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

function studentMsg(totalScores, studentScore) {
  let average = getAverage(totalScores);
  let grade = getGrade(studentScore);

  return `Class average: ${average}. Your grade: ${grade}. You ${
    hasPassingGrade(studentScore) ? "passed" : "failed"
  } the course.`;
}

```
