---
id: 662693f82c91a66be46c881b
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

Un profesor ha terminado de calificar los exámenes de sus estudiantes y necesita tu ayuda para calcular la puntuación media de la clase.

Completa la función `getAverage` que toma un array con las puntuaciones de las evaluaciones y devuelve la puntuación promedio.

El promedio se calcula sumando todas las puntuaciones y dividiendo por el número total de puntuaciones.

```md
average = sum of all scores / total number of scores
```

Se te han proporcionado un par de llamadas de funciones para que puedas probar tu código.

**Consejos**

- Puedes usar un bucle para iterar sobre el arreglo `puntuación` y sumar todas puntuaciones.
- Puedes usar la propiedad `length` para obtener el número total de puntuaciones.

# --hints--

Tu función `getAverage` debe devolver un número.

```js
assert.strictEqual(typeof getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 'number');
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` debe devolver `71.7`.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` debe devolver `85.4`.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` debe devolver `92.4`.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

Tu función `getAverage`, que calcula el promedio, debería devolver la puntuación media de las evaluaciones.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function getAverage(scores) {

}

console.log(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
console.log(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]));
--fcc-editable-region--
```
