---
id: 67ef8f695db52583424bd118
title: Розробіть сторінку з рецензією на фільм
challengeType: 25
dashedName: design-a-movie-review-page
demoType: onClick
---

# --description--

**Мета:** виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Ви повинні мати елемент `main`.
2. В межах елемента `main` має бути елемент `h1` з назвою фільму.
3. Під елементом `h1` має бути елемент `img` з постером фільму. Елемент `img` повинен мати атрибут `alt`, текст якого описує зображення. Якщо хочете, можете використати це зображення: `https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png`.
4. Ви повинні мати елемент `p`, який містить короткий опис фільму.
5. Ви повинні мати ще один елемент `p`, який містить оцінку фільму. Цей елемент повинен містити (в заданому порядку):
   - елемент `strong` з текстом `Movie Rating:`;
   - елемент `span` з атрибутом `aria-hidden` зі значенням `true`, що містить візуальне представлення оцінки у вигляді зірочок `⭐⭐⭐⭐⭐⭐⭐⭐⭐☆`;
   - числове значення оцінки в дужках (наприклад, `9.2/10`) після елемента `span`.
6. Ви повинні мати елемент `h2` з текстом `Cast Members`.
7. Ви повинні мати елемент `ul`.
8. В межах елемента `ul` має бути декілька елементів `li`, кожен з яких містить елемент `strong` з іменем актора та його роллю, розділеними словом `as` (наприклад, `James Holloway as Ethan Carter`).

# --hints--

Ви повинні мати елемент `main`.

```js
assert.isNotNull(document.querySelector("main"));
```

В межах елемента `main` має бути елемент `h1`, який містить назву фільму.

```js
const h1 = document.querySelector("main > h1");
assert.isAbove(h1?.innerText.length, 0);
```

Під елементом `h1` має бути елемент `img` з описовим атрибутом `alt`.

```js
const img = document.querySelector("main > h1 + img");
assert.isNotNull(img);
assert.isNotEmpty(img.getAttribute("alt"));
```

Після зображення має бути елемент `p`, який містить короткий опис фільму.

```js
const description = document.querySelector("main img + p");
assert.isNotNull(description);
assert.isNotEmpty(description.innerText);
```

Ви повинні мати ще один елемент `p`, який містить оцінку фільму.

```js
const paragraphs = document.querySelectorAll("main p");
assert.isAtLeast(paragraphs.length, 2);
```

В межах другого елемента `p` має бути елемент `strong` з текстом `Movie Rating:`.

```js
const strongEl = document.querySelector("main p:nth-of-type(2) strong");
assert.match(strongEl?.innerText.trim(), /^Movie Rating:?$/);
```

В межах абзацу з оцінкою має бути елемент `span`.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.isNotNull(spanEl);
```

В межах абзацу з оцінкою має бути елемент `span` з десятьма заповненими (⭐) або порожніми (☆) зірочками, а також числове значення в дужках після елемента `span`.

```js  
const spanEl = document.querySelector("main p:nth-of-type(2) span");
const fullRatingText = document.querySelector("main p:nth-of-type(2)")?.textContent.replace(/\s+/g, ' ').trim();

assert.match(spanEl?.innerText.trim(), /^[⭐☆]{10}$/);
assert.match(fullRatingText, /Movie Rating:?\s*[⭐☆]{10}\s*\(\s*\d+(?:\.\d+)?\s*\/\s*10\s*\)$/);
```

Елемент `span` в межах абзацу з оцінкою повинен мати атрибут `aria-hidden` зі значенням `true`.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.strictEqual(spanEl?.getAttribute("aria-hidden"), "true");
```

Ви повинні мати елемент `h2`.

```js
assert.exists(document.querySelector("h2"));
```

Після другого елемента `p` має бути елемент `h2`.

```js
const allParagraphs = document.querySelectorAll("p");
assert.isAtLeast(allParagraphs.length, 2);

const h2 = document.querySelector("h2");
const main = document.querySelector("main");

assert.exists(h2);
assert.exists(main);

const elements = Array.from(main.querySelectorAll("*"));
const lastPIndex = elements.indexOf(allParagraphs[1]);
const h2Index = elements.indexOf(h2);

assert.isAbove(h2Index, lastPIndex);
```

Елемент `h2` повинен мати текст `Cast Members`.

```js
const h2 = document.querySelector("h2");
assert.equal(h2?.textContent, "Cast Members");
```

Після елемента `h2` має бути елемент `ul`.

```js
const ul = document.querySelector("main h2 + ul");
assert.isNotNull(ul);
```

Елемент `ul` повинен містити декілька елементів `li`.

```js
const liItems = document.querySelectorAll("main ul li");
assert.isAbove(liItems.length, 1);
```

Кожен елемент `li` повинен містити елемент `strong` з іменем актора і його персонажем, розділеними словом `as` (наприклад, `James Holloway as Ethan Carter`).

```js
const liItems = document.querySelectorAll("main ul li");  
assert.isNotEmpty(liItems);
liItems.forEach(li => {
  const strong = li.querySelector("strong");
  assert.isNotNull(strong);
  const fullText = li.textContent.trim();
  const actorName = strong.textContent.trim();
  assert.isTrue(fullText.startsWith(actorName));
  const characterPart = fullText.slice(actorName.length).trim();
  assert.match(characterPart, /^as\s.+/);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
 <html lang="en">
 <head>
     <meta charset="UTF-8">
     <meta name="viewport" content="width=device-width, initial-scale=1.0">
     <title>Movie Review</title>
 </head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Movie Review</title>
  </head>
  <body>
    <main class="review">
      <h1>Movie Review: Rise Beyond</h1>
      <img
        src="https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png"
        alt="The Rise Beyond movie poster, starring Ethan Carter. A portrait of a young boy standing on a stage with his back toward the audience, holding a guitar at his side, staring at a backdrop of a cityscape with the sun rising behind its tall buildings."
      />
      <p>
        Rise Beyond is an uplifting drama that tells the story of Ethan Carter,
        a young man from a small town who dreams of becoming a world-class
        musician. Facing financial struggles and self-doubt, Ethan embarks on a
        journey of resilience, guided by his mentor Professor Adams and
        childhood friend Lena Mitchell. Through setbacks, failures, and moments
        of triumph, Rise Beyond reminds us that our greatest potential lies
        beyond our fears.
      </p>
      <p>
        <strong>Movie Rating:</strong>
        <span aria-hidden="true">⭐⭐⭐⭐⭐⭐⭐⭐⭐☆</span> (9.2/10)
      </p>

      <h2>Cast Members</h2>
      <ul>
        <li><strong>James Holloway</strong> as Ethan Carter</li>
        <li><strong>Olivia Sterling</strong> as Lena Mitchell</li>
        <li><strong>William Lancaster</strong> as Professor Adams</li>
        <li>
          <strong>Maria Collins</strong> as Evelyn Carter (Ethan’s Mother)
        </li>
      </ul>
    </main>
  </body>
</html>
```
