---
id: 6718d2d59337c822ecb697ff
title: Ein Programm zur Verwaltung von Bankkonten erstellen
challengeType: 26
dashedName: build-a-bank-account-management-program
saveSubmissionToDB: true
---

# --description--

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Klasse namens `BankAccount` mit einem Konstruktor definieren. Der Konstruktor sollte die Standard-`balance` auf `0` setzen und ein leeres Array namens `transactions` initialisieren, um Transaktionsdatensätze als Objekte zu speichern.

2. Jede in dem `transactions`-Array gespeicherte Transaktion sollte ein Objekt mit zwei Eigenschaften sein: `type` und `amount`. Die `type`-Eigenschaft sollte entweder `deposit` oder `withdraw` sein, und die `amount`-Eigenschaft sollte den eingezahlten oder abgehobenen Betrag enthalten.

3. Sie sollten eine Methode namens `deposit` definieren, die den Einzahlungsbetrag als Parameter erhält. Wenn der Einzahlungsbetrag größer als `0` ist, sollte sie:

   - Ein neues Objekt mit dem Typ `deposit` und dem eingezahlten Betrag zum Transaktionsarray hinzufügen.
   - Den Kontostand aktualisieren.
   - `"Successfully deposited $[amount]. New balance: $[balance]"` zurückgeben.

4. Wenn der Betrag kleiner oder gleich `0` ist, sollte die `deposit`-Methode `"Deposit amount must be greater than zero."` zurückgeben.

5. Sie sollten eine Methode namens `withdraw` definieren, die einen Betrag als Parameter erhält. Diese Methode sollte den aktuellen Kontostand entsprechend der Abhebungen aktualisieren. Wenn der abzuhebende Betrag größer als `0` und kleiner oder gleich dem aktuellen Kontostand ist, sollte sie:

   - Ein neues Objekt mit dem Typ `withdraw` und dem abgehobenen Betrag zum Transaktionsarray hinzufügen.
   - Den Kontostand aktualisieren.
   - `"Successfully withdrew $[amount]. New balance: $[balance]"` zurückgeben.

6. Wenn der abzuhebende Betrag kleiner oder gleich `0` oder größer als der aktuelle Kontostand ist, sollte die `withdraw`-Methode `"Insufficient balance or invalid amount."` zurückgeben.


7. Sie sollten eine Methode namens `checkBalance` definieren, die den aktuellen Kontostand im Format `"Current balance: $[balance]"` zurückgibt.

8. Sie sollten eine Methode namens `listAllDeposits` definieren, die durch das Transaktionsarray iteriert und alle Einzahlungen im Format `"Deposits: amount,amount,..."` zurückgibt.

9. Sie sollten eine Methode namens `listAllWithdrawals` definieren, die durch das Transaktionsarray iteriert und alle Abhebungen im Format `"Withdrawals: amount,amount,..."` zurückgibt.

10. Sie sollten eine neue Instanz von `BankAccount` mit dem Namen `myAccount` erstellen.

11. Ihr `myAccount`-Bankkonto sollte mindestens fünf Transaktionen haben.

12. Ihr `myAccount`-Bankkonto sollte mindestens zwei Einzahlungen haben.

13. Ihr `myAccount`-Bankkonto sollte mindestens zwei Abhebungen haben.

14. Ihr `myAccount`-Bankkonto sollte einen Kontostand größer als `$100` haben.

# --hints--

Sie sollten eine Klasse namens `BankAccount` definieren.

```js
assert.isFunction(BankAccount);
assert.isObject(new BankAccount());
```

Das `BankAccount`-Objekt sollte anfangs eine `balance` von `0` und ein leeres Array `transactions` zum Speichern von Transaktionsdatensätzen haben.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.balance, 0);
assert.equal(myAccount.transactions.length, 0);
```

Sie sollten eine `deposit`-Methode haben, die den Einzahlungsbetrag als Parameter erhält.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.deposit);
assert.lengthOf(myAccount.deposit, 1);
```

Sie sollten eine `withdraw`-Methode haben, die den Abhebungsbetrag als Parameter erhält.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.withdraw);
assert.lengthOf(myAccount.withdraw, 1);
```

Sie sollten eine `checkBalance`-Methode haben, die den aktuellen Kontostand überprüft.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.checkBalance);
```

Sie sollten eine `listAllDeposits`-Methode haben, die alle Einzahlungen auflistet.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllDeposits);
```

Sie sollten eine `listAllWithdrawals`-Methode haben, die alle Abhebungen auflistet.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllWithdrawals);
```

`BankAccount.deposit(100)` sollte `"Successfully deposited $100. New balance: $100"` zurückgeben.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(100), "Successfully deposited $100. New balance: $100")

// prevent hardcoding
assert.equal(myAccount.deposit(200), "Successfully deposited $200. New balance: $300")

```

`BankAccount.deposit(-50)` sollte `"Deposit amount must be greater than zero."` zurückgeben.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(-50), "Deposit amount must be greater than zero.")

// prevent hardcoding
assert.equal(myAccount.deposit(-30), "Deposit amount must be greater than zero.")

```

`BankAccount.deposit(0)` sollte `"Deposit amount must be greater than zero."` zurückgeben.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(0), "Deposit amount must be greater than zero.")
```

Wenn der Kontostand `100` ist, sollte `BankAccount.withdraw(150)` `"Insufficient balance or invalid amount."` zurückgeben.

```js
const myAccount = new BankAccount();
myAccount.deposit(100);
assert.equal(myAccount.withdraw(150), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(250), "Insufficient balance or invalid amount.")

```

`BankAccount.withdraw(-50)` sollte `"Insufficient balance or invalid amount."` zurückgeben.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(-50), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(-30), "Insufficient balance or invalid amount.")
```

`BankAccount.withdraw(0)` sollte `"Insufficient balance or invalid amount."` zurückgeben.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(0), "Insufficient balance or invalid amount.")
```

Wenn der Kontostand `200` ist, sollte `BankAccount.withdraw(150)` `"Successfully withdrew $150. New balance: $50"` zurückgeben.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.withdraw(150), "Successfully withdrew $150. New balance: $50")
```

Wenn der Kontostand `200` ist, sollte `BankAccount.checkBalance()` `"Current balance: $200"` zurückgeben.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.checkBalance(), "Current balance: $200")

// prevent hardcoding
myAccount.deposit(300);
assert.equal(myAccount.checkBalance(), "Current balance: $500")
```

Wenn Sie `10`, `35`, `90` einzahlen, sollte die `listAllDeposits`-Methode `"Deposits: 10,35,90"` zurückgeben.

```js
const myAccount = new BankAccount();
myAccount.deposit(10);
myAccount.deposit(35);
myAccount.deposit(90);
assert.equal(myAccount.listAllDeposits(), "Deposits: 10,35,90")
```

Wenn Sie `20`, `50`, `100` abheben, sollte die `listAllWithdrawals`-Methode `"Withdrawals: 20,50,100"` zurückgeben.

```js
const myAccount = new BankAccount();
myAccount.deposit(500);

myAccount.withdraw(20);
myAccount.withdraw(50);
myAccount.withdraw(100);
assert.equal(myAccount.listAllWithdrawals(), "Withdrawals: 20,50,100")
```

Sie sollten eine Instanz von `BankAccount` mit dem Namen `myAccount` haben.

```js
assert.exists(myAccount);
assert.instanceOf(myAccount, BankAccount);
```

Ihr `myAccount`-Bankkonto sollte mindestens fünf Transaktionen haben.

```js
assert.isAtLeast(myAccount.transactions.length, 5);
```

Ihr `myAccount`-Bankkonto sollte mindestens zwei Einzahlungen haben.

```js
const depositCount = myAccount.transactions.filter(transaction => transaction.type === "deposit").length;
assert.isAtLeast(depositCount, 2);
```

Ihr `myAccount`-Bankkonto sollte mindestens zwei Abhebungen haben.

```js
const withdrawalCount = myAccount.transactions.filter(transaction => transaction.type === "withdraw").length;
assert.isAtLeast(withdrawalCount, 2);
```


Ihr `myAccount`-Bankkonto sollte einen Kontostand größer als `$100` haben.

```js
assert.isAbove(myAccount.balance, 100);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class BankAccount {
    constructor() {
        this.balance = 0;
        this.transactions = [];
    }

    deposit(amount) {
        if (amount > 0) {
            this.transactions.push({ type: 'deposit', amount });
            this.balance += amount;
            return `Successfully deposited $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Deposit amount must be greater than zero.';
        }
    }

    withdraw(amount) {
        if (amount > 0 && amount <= this.balance) {
            this.transactions.push({ type: 'withdraw', amount });
            this.balance -= amount;
            return `Successfully withdrew $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Insufficient balance or invalid amount.';
        }
    }

    checkBalance() {
        return `Current balance: $${this.balance}`;
    }

    listAllDeposits() {
        const deposits = this.transactions
            .filter(transaction => transaction.type === 'deposit')
            .map(transaction => transaction.amount);
        return `Deposits: ${deposits}`;
    }

    listAllWithdrawals() {
        const withdrawals = this.transactions
            .filter(transaction => transaction.type === 'withdraw')
            .map(transaction => transaction.amount);
        return `Withdrawals: ${withdrawals}`;
    }
}

const myAccount = new BankAccount();
  
console.log(myAccount.deposit(0)); 
console.log(myAccount.deposit(20)); 
console.log(myAccount.deposit(500));
console.log(myAccount.withdraw(30));
console.log(myAccount.withdraw(55));
console.log(myAccount.withdraw(100));
console.log(myAccount.checkBalance());
console.log(myAccount.listAllDeposits());
console.log(myAccount.listAllWithdrawals());
```
