---
id: 684ccbdc1ef89e042b744e09
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

A próxima parte do workshop é adicionar o botão de fechar. 

Abaixo do seu elemento `p`, adicione um elemento `button` com um `aria-label` definido como `"Dismiss"`.Dentro desse elemento `button`, adicione um elemento `span` com um atributo `aria-hidden`.O texto do elemento `span` deve ser `&times;`. 

*NOTA*: Como este workshop será focado apenas em trabalhar com estilos Tailwind CSS, o `button` não terá nenhum tipo de interatividade adicionada a ele. 

# --hints--

Você deve ter um elemento `button`.

```js
assert.exists(document.querySelector("button"));
```
 
Seu elemento `button` deve ter um atributo `aria-label`.

```js
const buttonEl = document.querySelector("button");
assert.isTrue(buttonEl?.hasAttribute("aria-label"));
```

Seu elemento `button` deve ter um atributo `aria-label` definido como `"Dismiss"`.

```js
const buttonEl = document.querySelector("button");
assert.equal(buttonEl?.getAttribute("aria-label"), "Dismiss");
```

Seu elemento `button` deve ter um elemento `span` aninhado dentro.

```js
assert.exists(document.querySelector("button span"));
```

Seu elemento `span` deve ter um atributo `aria-hidden`.

```js
const spanEl = document.querySelector("span");
assert.isTrue(spanEl?.hasAttribute("aria-hidden"));
```

Seu elemento `span` deve conter o texto `&times;`.

```js
const spanEl = document.querySelector("span");
assert.equal(spanEl?.innerHTML, "×");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--

      --fcc-editable-region--
    </div>
  </body>
</html>
```
