---
id: 684cff0bf0a4ca0ad7e42990
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

O próximo passo seria arredondar os cantos da borda. Para definir a propriedade `border-radius` usando Tailwind CSS, você pode usar uma classe utilitária com o seguinte padrão: `rounded-size`, onde `size` representa o raio. 

Aqui está um exemplo de como aplicar um raio de borda extra-grande a um elemento:

```html
<p class="rounded-xl">This is border radius.</p>
```

Dentro da sua tag de abertura `div`, aplique um border radius com o tamanho `md`. 

# --hints--

Sua tag de abertura `div` deve ter a classe utilitária `rounded-md`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("rounded-md"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
