---
id: 684cca09aac0fc035db4099f
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro do seu elemento `p`, adicione o texto `Error!` Algo deu errado. Por favor, tente novamente.` onde `Error!` deve estar envolvido dentro de um elemento `strong`.

# --hints--

Você deve ter um elemento `strong` dentro do seu elemento `p`.

```js
assert.exists(document.querySelector("p strong"));
```

Seu elemento `strong` deve conter o texto `Error!`.

```js
assert.equal(document.querySelector("strong")?.textContent, "Error!");
```

Seu elemento `p` deve ter o texto correto que termina com `Something went wrong. Please try again.`

```js
const remainingText = "Something went wrong. Please try again."
const pElement = document.querySelector("p");
assert.isTrue(pElement?.textContent.trim().endsWith(remainingText));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      --fcc-editable-region--
      <p>

      </p>
      --fcc-editable-region--
    </div>
  </body>
</html>
```
