---
id: 684cd8bf5445de089af9c837
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Agora é hora de aplicar uma cor de borda. Para definir a propriedade `border-color` usando Tailwind CSS, você pode usar uma classe utilitária com o seguinte padrão: `border-color-number`, onde `number` representa o tom para a cor da borda. 

Aqui está um exemplo de como aplicar uma cor de borda roxa a um elemento:

```html
<p class="border-purple-400">This is border-color.</p>
```

Dentro da sua tag de abertura `div`, aplique uma cor de borda vermelha com um tom `300`. 

# --hints--

Sua tag de abertura `div` deve ter a classe utilitária `border-red-300`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-red-300"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
