---
id: 684cc59eed2706fa50e8933c
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você aprenderá como começar a trabalhar com o básico do Tailwind CSS construindo um componente de mensagem de erro. 

Nos primeiros passos, você vai se concentrar em construir a marcação HTML inicial. 

Dentro do elemento `body`, adicione um elemento `div`.Então dentro desse elemento `div`, adicione um elemento `p`.

# --hints--

Você deve ter um elemento `div`.

```js
assert.exists(document.querySelector("div"));
```

Você deve ter um elemento `p`.

```js
assert.exists(document.querySelector("p"));
```

Seu elemento `p` deve estar dentro do elemento `div`.

```js
assert.exists(document.querySelector("div p"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
