---
id: 684cd59ec94b5e063bf776f0
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Seria bom ter uma borda ao redor daquele fundo vermelho claro. 

Para definir o `border-width` no Tailwind CSS, você pode usar a classe utilitária `border-number` onde `number` representará a leveza ou espessura da sua borda. 

Aqui está um exemplo de aplicação de um `border-width` de `4px` em todos os lados:

```html
<p class="border-4">This is how to set the border-width.</p>
```

Dentro da sua tag de abertura `div`, adicione a classe utilitária correta para um `border-width` definido como `2px` em todos os lados.

# --hints--

Sua tag de abertura `div` deve ter uma classe `border-2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
