---
id: 684d1942a1904a14a3709c6c
title: Passo 20
challengeType: 0
dashedName: step-20
---

# --description--

A estilização do seu componente de mensagem de erro está quase completa, mas ainda há uma coisa para corrigir.Como você definiu tamanhos de fonte diferentes para o parágrafo e o texto do botão, o alinhamento vertical entre esses itens flex está incorreto. 

Para posicionar itens flex ao longo do eixo cruzado de um container usando Tailwind CSS, você pode usar o seguinte padrão: `items-value`. `value` pode ser qualquer coisa entre `start`, `end`, `center`, `baseline` e mais. 

Dentro da sua tag de abertura `div`, use a classe utilitária correta para definir a propriedade `align-items` como `center`.

E com essa última alteração, seu componente de mensagem de erro está completo!

# --hints--

Sua tag de abertura `div` deve ter a classe utilitária `items-center`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
    --fcc-editable-region--
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto flex gap-4 justify-center items-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
