---
id: 684d11ffd15a531225d081e9
title: Passo 17
challengeType: 0
dashedName: step-17
---

# --description--

Agora que os estilos foram aplicados ao elemento `div`, é hora de aplicar alguns estilos ao elemento `p`. 

Atualmente, a cor do texto para o elemento `paragraph` é preta.Mas seria bom se a cor fosse um vermelho escuro em vez disso.

Em uma etapa anterior, você aprendeu como trabalhar com cores de fundo assim: `bg-red-100`.Você pode fazer algo semelhante com a cor do texto seguindo este padrão: `text-color-shade`. 

Aqui está um exemplo de como aplicar uma cor de texto azul a um elemento:

```html
<p class="text-blue-500">This has blue text.</p>
```

Comece adicionando um atributo class ao seu elemento `p`.Então, aplique a classe utilitária correta para definir a cor do texto vermelho com o tom definido para `700`.

# --hints--

Seu elemento `p` deve ter um atributo class.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.hasAttribute("class"));
```

Sua tag de abertura `p` deve ter a classe utilitária `text-red-700`.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.classList.contains("text-red-700"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p>
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
