---
id: 684d0d369e2b3b0f288d2c20
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

Seria bom centralizar horizontalmente o elemento `div`.Para aplicar `margin` em dois lados ao mesmo tempo, você pode usar os seguintes padrões: `mx-number` ou `my-number`. 

`mx` define a margem para a esquerda e direita, enquanto `my` define a margem para o topo e a parte inferior.

Aqui está um exemplo de aplicação das margens superior e inferior a um elemento:

```html
<p class="my-8">This is margin.</p>
```

Dentro da sua tag de abertura `div`, use a classe utilitária correta para definir a margem esquerda e direita como `auto`.

# --hints--

Sua tag de abertura `div` deve ter a classe utilitária `mx-auto`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
