---
id: 684d07cbfa43a20d6a002ed7
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Os estilos para o elemento `div` estão começando a se formar. Mas seria bom se houvesse algum espaço ao redor do texto. 

Para aplicar `padding` a um elemento usando Tailwind CSS, você pode usar o seguinte padrão: `p-number`, onde `number` representa os pixels que você deseja aplicar em todos os quatro lados. 

Aqui está um exemplo de aplicação de `padding` de `8px` a um elemento:

```html
<p class="p-8">This is padding.</p>
```

Dentro da sua tag de abertura `div`, use a classe utilitária correta para aplicar `padding` de `4px`.

# --hints--

Sua tag de abertura `div` deve ter a classe utilitária `p-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
