---
id: 684d0eeb9c81e5109e52fe8b
title: Passo 15
challengeType: 0
dashedName: step-15
---

# --description--

O elemento `div` está começando a ficar melhor.Mas agora o `button` de fechar está um pouco muito próximo do elemento parágrafo. 

Para criar espaço entre os itens flex, você pode usar o seguinte padrão: `gap-number`. `number` representa o espaço fixo entre as linhas e colunas em um layout flex.

Aqui está um exemplo de aplicação de um `gap` a um elemento:

```html
<p class="flex gap-8">This is gap.</p>
```

Dentro da sua tag de abertura `div`, use a classe utilitária correta para definir o `gap` como `4`.

# --hints--

Sua tag de abertura `div` deve ter a classe utilitária `gap-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("gap-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
