---
id: 684d09cad4ec900e41db2e74
title: Passo 11
challengeType: 0
dashedName: step-11
---

# --description--

Neste momento o elemento `div` está posicionado no canto superior esquerdo da tela de visualização.Mas seria bom adicionar um pouco de espaço no topo do elemento.Para aplicar `margin` a um elemento usando Tailwind CSS, você pode usar o seguinte padrão: `m-number`, onde `number` representa os pixels fixos a serem aplicados em todos os quatro lados.

Se você quiser aplicar margem a apenas um lado, pode usar qualquer um dos seguintes padrões: `mt-number` (margin top), `mb-number` (margin bottom), `mr-number` (margin right) e `ml-number` (margin left).

Aqui está um exemplo de como aplicar `margin-bottom` a um elemento:

```html
<p class="mb-8">This is margin.</p>
```

Dentro da sua tag de abertura `div`, use a classe utilitária correta para aplicar `margin-top` de `4px`.

# --hints--

Sua tag de abertura `div` deve ter a classe utilitária `mt-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mt-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
