---
id: 684d160a2370b01310ae0fda
title: Passo 18
challengeType: 0
dashedName: step-18
---

# --description--

A nova cor vermelha fica boa, mas seria melhor se o texto fosse maior. Para definir o `font-size` de um elemento usando Tailwind CSS, você pode usar o seguinte padrão: `text-value`. 

Aqui está um exemplo de como definir o `font-size` para um elemento como `0.875rem` ou (`14px`):

```html
<p class="text-sm">This is for font-size.</p>
```

Dentro do seu elemento de abertura `p`, use a classe utilitária correta para definir o `font-size` para `xl` que representa `1.25rem` ou (`20px`).

# --hints--

Sua tag de abertura `p` deve ter a classe utilitária `text-xl`.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl.classList.contains("text-xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p class="text-red-700">
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
