---
id: 684d0e11094c3e0fe1fb98f5
title: Passo 14
challengeType: 0
dashedName: step-14
---

# --description--

No momento, o elemento de parágrafo está empilhado sobre o elemento de `button` de dispensa.Mas seria bom que o elemento `button` estivesse na mesma linha que o parágrafo.É aqui que o Flexbox entra em cena.

Para definir o elemento `div` com um layout Flexbox no Tailwind CSS, você pode usar a classe utilitária `flex`.

# --hints--

Sua tag de abertura `div` deve ter a classe utilitária `flex`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("flex"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
