---
id: 684d105e488dbd1166a717ae
title: Passo 16
challengeType: 0
dashedName: step-16
---

# --description--

Se você aumentar o tamanho da janela de visualização, notará que todos os itens ainda estão alinhados à esquerda no contêiner. Seria bom se os itens flex fossem centralizados. 

Para posicionar itens flex ao longo do eixo principal de um container com Tailwind CSS, você pode usar o seguinte padrão: `justify-value`. `value` pode ser `start`, `end`, `center`, `between` e mais. Qualquer valor válido usado no CSS puro com `justify-content` pode ser usado aqui com Tailwind CSS. 

Dentro da sua tag de abertura `div`, use a classe utilitária correta para definir a propriedade `justify-content` como `center`.

# --hints--

Sua tag de abertura `div` deve ter a classe utilitária `justify-center`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("justify-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
