---
id: 684d009994a3840b98c9aa3a
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

No momento, a largura do elemento `div` abrange toda a largura da página. Para aplicar larguras no Tailwind CSS, você tem algumas opções, como usar o padrão `w-number` para larguras fixas ou o padrão `w-fraction` para porcentagens.  

Aqui está um exemplo de aplicação de uma largura definida para `66%` do seu contêiner pai:

```html
<p class="w-2/3">This is a percentage width.</p>
```

Dentro da sua tag de abertura `div`, aplique uma largura definida para `50%` do seu contêiner pai.

# --hints--

Sua tag de abertura `div` deve ter a classe utilitária `w-1/2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
