---
id: 68501d32b69f9d7756334ffc
title: Passo 13
challengeType: 0
dashedName: step-13
---

# --description--

Se você ampliar a janela de visualização, deverá ver o elemento `div` centralizado horizontalmente na página. Mas para dispositivos menores, a largura do elemento `div` abrange toda a largura da página.  

Atualize a classe `mx-auto` para que ela seja aplicada apenas em dispositivos `md` e maiores.

# --hints--

Sua tag de abertura `div` deve ter a classe utilitária `md:mx-auto`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:mx-auto"));
```

Você não deve mais ter a classe `mx-auto` na sua tag `div` de abertura.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
