---
id: 66c8ba975ee7230e29f6c4ad
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

للتحقق من القرض الذي يحق للمستخدم الحصول عليه بناءً على `annualIncome` و`creditScore`، عليك استخدام جملة `if/else if` أو ثلاثي داخل دالة `getLoanMessage`. ثم تُعيد الرسالة المناسبة في كتلة كل شرط.

ابدأ بقرض الدوبلكس، تحقق مما إذا كان `annualIncome` أكبر من أو يساوي `minIncomeForDuplex` وأيضًا إذا كان `creditScore` أكبر من أو يساوي `minCreditScoreForDuplex`.

إذا تحقق هذا الشرط، يكون المتقدم مؤهلاً لقرض الدوبلكس والقروض الأخرى. لذا، داخل التحقق، أعد السلسلة النصية `"You qualify for a duplex, condo, and car loan."`

# --hints--

يجب أن تُعيد دالة `getLoanMessage` سلسلة نصية.

```js
assert.isString(getLoanMessage(65000, 750));
assert.isString(getLoanMessage(60000, 750));
assert.isString(getLoanMessage(65000, 700));
assert.isString(getLoanMessage(60000, 700));
```

يجب أن تُعيد دالة `getLoanMessage` السلسلة النصية `"You qualify for a duplex, condo, and car loan."`.

```js
assert.strictEqual(getLoanMessage(65000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(65000, 700), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 700), "You qualify for a duplex, condo, and car loan.");
```

يجب أن تُعيد دالة `getLoanMessage` القيمة `undefined` إذا لم يستوفِ الدخل السنوي ودرجة الائتمان للمتقدم متطلبات قرض الدوبلكس.

```js
assert.isUndefined(getLoanMessage(59000, 690));
assert.isUndefined(getLoanMessage(59000, 700));
assert.isUndefined(getLoanMessage(59000, 750));
assert.isUndefined(getLoanMessage(60000, 690));
assert.isUndefined(getLoanMessage(65000, 690));
```
 
# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
