---
id: 66c8ba975ee7230e29f6c4b2
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

حان الوقت الآن لاختبار دالة `getLoanMessage` الخاصة بك.

استخدم الجدول أدناه لإنشاء `4` متغيرات وقيمها:

| اسم المتغير | القيمة   |
| ----------- | ------- |
| `duplexLoanMsg`       | `getLoanMessage(85000, 850)` |
| `condoLoanMsg`   | `getLoanMessage(65000, 690)`   |
| `carLoanMsg`   | `getLoanMessage(45000, 660)`   |
| `noLoanMsg`   | `getLoanMessage(25000, 550)`   |

بعد ذلك، سجّل كل متغير في وحدة التحكم لرؤية الرسائل.

بهذا، يكتمل مشروع فاحص تأهيل القرض الخاص بك!

# --hints--

يجب أن يكون لديك متغير اسمه `duplexLoanMsg`.

```js
assert.isNotNull(duplexLoanMsg);
```

يجب تعيين متغير `duplexLoanMsg` إلى نتيجة `getLoanMessage(85000, 850)`.

```js
assert.equal(duplexLoanMsg, getLoanMessage(85000, 850));
```

يجب أن تسجل متغير `duplexLoanMsg` في وحدة التحكم.

```js
assert.match(code, /console\.log\(duplexLoanMsg\)/)
```

يجب أن يكون لديك متغير اسمه `condoLoanMsg`.

```js
assert.isNotNull(condoLoanMsg);
```

يجب تعيين متغير `condoLoanMsg` إلى نتيجة `getLoanMessage(65000, 690)`.

```js
assert.equal(condoLoanMsg, getLoanMessage(65000, 690));
```

يجب أن تسجل متغير `condoLoanMsg` في وحدة التحكم.

```js
assert.match(code, /console\.log\(condoLoanMsg\)/)
```

يجب أن يكون لديك متغير اسمه `carLoanMsg`.

```js
assert.isNotNull(carLoanMsg);
```

يجب تعيين متغير `carLoanMsg` إلى نتيجة `getLoanMessage(45000, 660)`.

```js
assert.equal(carLoanMsg, getLoanMessage(45000, 660));
```

يجب أن تسجل متغير `carLoanMsg` في وحدة التحكم.

```js
assert.match(code, /console\.log\(carLoanMsg\)/)
```

يجب أن يكون لديك متغير اسمه `noLoanMsg`.

```js
assert.isNotNull(noLoanMsg);
```

يجب تعيين متغير `noLoanMsg` إلى نتيجة `getLoanMessage(25000, 550)`.

```js
assert.equal(noLoanMsg, getLoanMessage(25000, 550));
```

يجب أن تسجل متغير `noLoanMsg` في وحدة التحكم.

```js
assert.match(code, /console\.log\(noLoanMsg\)/)
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

const duplexLoanMsg = getLoanMessage(85000, 850)
console.log(duplexLoanMsg)

const condoLoanMsg = getLoanMessage(65000, 690)
console.log(condoLoanMsg)

const carLoanMsg = getLoanMessage(45000, 660)
console.log(carLoanMsg)

const noLoanMsg = getLoanMessage(25000, 550)
console.log(noLoanMsg)
```
