---
id: 66c8ba975ee7230e29f6c4b0
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

الآن، يجب أن تتحقق مما إذا كان المتقدم مؤهلًا لقرض سيارة فقط. إذا كان كذلك، أرجع السلسلة النصية `"You qualify for a car loan."`.

# --hints--

يجب أن تُرجع دالة `getLoanMessage` سلسلة نصية إذا كان المتقدم مؤهلًا لقرض سيارة.

```js
assert.isString(getLoanMessage(30000, 650));
```

يجب أن تُرجع دالة `getLoanMessage` السلسلة النصية `"You qualify for a car loan."`.

```js
assert.strictEqual(getLoanMessage(30000, 650), "You qualify for a car loan.");
```

يجب أن تُرجع دالة `getLoanMessage` القيمة `undefined` إذا لم يكن دخل المتقدم السنوي ودرجة الائتمان يلبّيان متطلبات قرض السيارة.

```js
assert.strictEqual(getLoanMessage(30000, 600), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
--fcc-editable-region--
  } 
  
--fcc-editable-region--
}
```
