---
id: 66c8ba41a77db20a93f9d7a1
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستراجع الشروط في JavaScript من خلال بناء تطبيق لفحص أهلية القرض.

سيتحقق التطبيق مما إذا كان المستخدم مؤهلاً لقرض دوبلكس، سيارة، وشقة بناءً على دخله السنوي ودرجة الائتمان الخاصة به. درجة الائتمان هي رقم يمثل مدى مهارتك في إدارة الأموال المقترضة.

للبدء، أنشئ المتغيرات والقيم التالية.

| اسم المتغير | القيمة   |
| ----------- | ------- |
| `minIncomeForDuplex`       | `60000` |
| `minCreditScoreForDuplex`   | `700`   |
| `minIncomeForCondo`   | `45000`   |
| `minCreditScoreForCondo`    | `680`   |
| `minIncomeForCar`   | `30000`   |
| `minCreditScoreForCar`   | `650`   |

# --hints--

يجب أن تنشئ متغير `minIncomeForDuplex` بقيمة `60000`.

```js
assert.equal(minIncomeForDuplex, 60000);
```

يجب أن تنشئ متغير `minCreditScoreForDuplex` بقيمة `700`.

```js
assert.equal(minCreditScoreForDuplex, 700);
```

يجب أن تنشئ متغير `minIncomeForCar` بقيمة `30000`.

```js
assert.equal(minIncomeForCar, 30000);
```

يجب أن تنشئ متغير `minCreditScoreForCar` بقيمة `650`.

```js
assert.equal(minCreditScoreForCar, 650);
```

يجب أن تنشئ متغير `minIncomeForCondo` بقيمة `45000`.

```js
assert.equal(minIncomeForCondo, 45000);
```

يجب أن تنشئ متغير `minCreditScoreForCondo` بقيمة `680`.

```js
assert.equal(minCreditScoreForCondo, 680);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
