---
id: 66c8ba975ee7230e29f6c4ac
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

عندما يكون المستخدم مؤهلاً للحصول على قرض، سترغب في عرض رسالة له في وحدة التحكم.

لهذا، ستبني دالة تحتوي على بعض الفحوصات التي تُرجع نوع القرض الذي يكون المتقدم مؤهلاً له.

أنشئ دالة فارغة باسم `getLoanMessage` تحتوي على معلمتين هما `annualIncome` و `creditScore`.

# --hints--

يجب أن تنشئ دالة `getLoanMessage`

```js
assert.isFunction(getLoanMessage)
```

يجب أن تحتوي دالة `getLoanMessage` على معلمتين هما `annualIncome` و `creditScore`.

```js
assert.match(getLoanMessage?.toString(), /\(\s*annualIncome,\s*creditScore\s*\)|\(\s*creditScore,\s*annualIncome\s*\)/);
```

يجب أن تكون دالة `getLoanMessage` فارغة.

```js
assert.match(getLoanMessage?.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

--fcc-editable-region--

--fcc-editable-region--
```
