---
id: 66c8ba975ee7230e29f6c4af
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

إذا كان الدخل السنوي للمتقدم أكبر من أو يساوي `minIncomeForCondo`، وأيضًا إذا كانت درجة الائتمان الخاصة به أكبر من أو تساوي `minCreditScoreForCondo`، فهو مؤهل للحصول على قرض شقة وسيارة.

تحقق مما إذا كان ذلك صحيحًا في الدالة `getLoanMessage`. إذا كان كذلك، أرجع السلسلة النصية `"You qualify for a condo and car loan."`

# --hints--

يجب أن تُرجع دالة `getLoanMessage` سلسلة نصية إذا كان المتقدم مؤهلًا للحصول على قرض شقة وسيارة.

```js
assert.isString(getLoanMessage(45000, 680));
```

يجب أن تُرجع دالة `getLoanMessage` السلسلة النصية `"You qualify for a condo and car loan."`.

```js
assert.strictEqual(getLoanMessage(45000, 680), "You qualify for a condo and car loan.");
```

يجب أن تُرجع دالة `getLoanMessage` القيمة `undefined` إذا لم يستوفِ الدخل السنوي ودرجة الائتمان للمتقدم متطلبات قرض الشقة.

```js
assert.strictEqual(getLoanMessage(45000, 650), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
