---
id: 656470d517833a39bb8b5608
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

El módulo `random` contiene un generador de números pseudoaleatorios. La mayoría de sus funcionalidades dependen de la función `random()`, que devuelve un número de punto flotante en el rango entre `0.0` (inclusive) y `1.0` (exclusivo).

Llama a la función `random()` del módulo `random` e imprime el resultado.

# --hints--

Debes imprimir el resultado de llamar a `random.random()`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.random\s*\(\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--

--fcc-editable-region--
```
