---
id: 656469e62377ab34a5bcb8d1
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Declara dos nuevas variables llamadas `digits` y `symbols` y asígnales `string.digits` y `string.punctuation`, respectivamente.

# --hints--

Debes declarar una variable llamada `digits`.

```js
({ test: () => assert(__userGlobals.has("digits")) })
```

Debes asignar `string.digits` a tu variable `digits`.

```js
({ test: () => assert(runPython(`
    import string
    digits == string.digits
  `))
})
```

Debes declarar una variable llamada `symbols`.

```js
({ test: () => assert(__userGlobals.has("symbols")) })
```

Debes asignar `string.punctuation` a tu variable `symbols`.

```js
({ test: () => assert(runPython(`
    import string
    symbols == string.punctuation
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string

letters = string.ascii_letters
--fcc-editable-region--
```
