---
id: 6564f0279e23ce924eedd1b2
title: Paso 51
challengeType: 20
dashedName: step-51
---

# --description--

Dado que el carácter de subrayado es un carácter válido para nombres de variables, está incluido en la clase de caracteres `\w` (equivalente a `[a-zA-Z0-9_]`), que puede usarse convenientemente para coincidir con nombres de variables.

Por lo tanto, la clase de caracteres `\W` es equivalente a `[^a-zA-Z0-9_]` con el carácter de guion bajo que no se coincide. Por esta razón no puedes usarla para coincidir con todos tus caracteres especiales.

Elimina las últimas tres líneas en tu código.

# --hints--

Debes eliminar las últimas tres líneas en tu código.

```js
({ test: () => {
  const commentless_code = __helpers.python.removeComments(code);
  assert.notMatch(commentless_code, /pattern\s*=\s*r("|')\\W\1/);
  assert.notMatch(commentless_code, /quote\s*=\s*("|')_\1/);
  assert.notMatch(commentless_code, /print\(\s*re\.findall\(\s*pattern\s*,\s*quote\s*\)\s*\)/);
}})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'\W')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = r'\W'
quote = '_'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
