---
id: 6564b8c9349bd76dc037967b
title: Paso 29
challengeType: 20
dashedName: step-29
---

# --description--

La función `search()` del módulo `re` analiza la cadena pasada como argumento buscando el primer lugar donde el patrón regex coincide con la cadena.

Declara una variable llamada `quote` y asigna la cadena `'Not all those who wander are lost.'` a esta variable. Luego, imprime el resultado de `pattern.search(quote)`.

# --hints--

Debes tener una variable `quote`.

```js
({ test: () => assert(__userGlobals.has("quote")) })
```

Debes asignar la cadena proporcionada a tu nueva variable `quote`.

```js
({ test: () => assert.equal(__userGlobals.get("quote"), "Not all those who wander are lost.") })
```

Debes imprimir `pattern.search(quote)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*pattern\.search\s*\(\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('i')

--fcc-editable-region--
```
