---
id: 656477845006313fbfea0ad1
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

Aunque el efecto pueda parecer igual a `random.choice()`, `secrets` te garantiza la fuente de aleatoriedad más segura que tu sistema operativo puede ofrecer.

Ahora, elimina tus dos llamadas a `print()`.

# --hints--

Debes eliminar tus dos llamadas a `print()`.

```js
({ test: () => assert.isFalse( /^print/m.test(code)) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols
--fcc-editable-region--
print(all_characters)
print(secrets.choice(all_characters))
--fcc-editable-region--
```
