---
id: 6564c67db26c417561ab510d
title: Paso 31
challengeType: 20
dashedName: step-31
---

# --description--

Como puedes ver en la salida, ahora tu regex coincide con la primera `l` dentro de la cadena.

En tu patrón, puedes agregar un cuantificador después de un carácter para especificar cuántas veces debe repetirse ese carácter. Por ejemplo, el cuantificador `+` significa que debe repetirse una o más veces.

Agrega un cuantificador `+` a tu patrón.

# --hints--

Debes modificar tu variable `pattern` a `re.compile('l+')`.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('l+')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('l')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
