---
id: 656474314bf37d3c83a85143
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

La función `choice()` del módulo `random` toma una secuencia y devuelve un elemento aleatorio de la secuencia.

Modifica tu llamada a `print()` para usar la función `choice()` y pasar `all_characters` como argumento.

# --hints--

Debes modificar tu llamada existente a `print()` para imprimir `random.choice(all_characters)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--
print(random.random())
--fcc-editable-region--
```
