---
id: 6564cbb6311a40783b3f5de6
title: Paso 36
challengeType: 20
dashedName: step-36
---

# --description--

Las clases de caracteres también te permiten indicar un rango de caracteres para coincidir. Necesitas especificar el carácter inicial y el carácter final separados por un guion, `-`. Los caracteres siguen el orden Unicode.

Modifica tu variable `pattern` para que coincida con cualquier letra `t` precedida por una letra minúscula en la variable `quote`. Usa el rango de caracteres de `a` a `z` para eso.

# --hints--

Debes modificar tu variable `pattern` para que coincida con cualquier letra `t` en `quote` precedida por una letra minúscula usando la clase `[a-z]`.

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "[a-z]t") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0123456789]')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'w[ha]'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
