---
id: 656475bbf1c2573de1d2c69c
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Cada vez que el código se ejecute, deberías ver un carácter aleatorio de la cadena `all_characters`. Esto es exactamente lo que quieres lograr para crear una contraseña aleatoria.

Sin embargo, el algoritmo en el que se basa `random` hace que los números pseudoaleatorios generados sean predecibles. Por lo tanto, aunque el módulo `random` es adecuado para las aplicaciones más comunes, no puede usarse para fines criptográficos, debido a su naturaleza determinista. 

En lugar de importar `random`, importa el módulo `secrets`. Luego cambia la llamada a `print()` para usar `secrets.choice(all_characters)`.

# --hints--

Debes importar el módulo `secrets` en lugar del módulo `random`.

```js
({
  test: () => {
    assert.match(code, /^import\s+secrets/m)
    assert.isFalse( /^import\s+random/m.test(code))
  }
})
```

Debes modificar tu llamada existente a `print()` para imprimir `secrets.choice(all_characters)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*secrets\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
print(random.choice(all_characters))
--fcc-editable-region--
```
