---
id: 6564d97a5ef99783216a6229
title: Paso 44
challengeType: 20
dashedName: step-44
---

# --description--

Python proporciona un tipo particular de cadena llamada cadena *raw*. Las cadenas raw tienen como prefijo un `r`. La diferencia clave con las cadenas normales radica en cómo manejan el carácter de barra invertida: en las cadenas raw, las barras invertidas se tratan como caracteres literales en lugar de caracteres de escape. Al escribir expresiones regulares, usar cadenas raw es una buena práctica, ya que generalmente pueden contener muchos caracteres `\`.

Convierte tu cadena `pattern` en una cadena raw anteponiéndole una `r`.

# --hints--

Debes modificar tu variable `pattern` a `r'\.'`.

```js
({ test: () => assert.match(code, /^pattern\s*=\s*r("|')\\\.\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '\.'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
