---
id: 6564dd65c3c2fa873a83d213
title: Paso 47
challengeType: 20
dashedName: step-47
---

# --description--

En una clase de caracteres, puedes combinar múltiples rangos escribiendo un rango tras otro dentro de los corchetes (sin ningún carácter adicional). Por ejemplo: `[a-d3-6]` es la combinación de `[a-d]` y `[3-6]`.

Ahora, modifica el último patrón regex para que coincida con cualquier carácter no alfanumérico. Combina los rangos `a-z`, `A-Z` y `0-9` en una sola clase de caracteres y añade un `^` como primer carácter para negar el patrón.

# --hints--

Debes completar tu cuarto patrón para que coincida con cualquier carácter no alfanumérico.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*r("|')\\d\1\s*\)\s*,\s*\(\s*lowercase\s*,\s*r("|')\[a-z\]\2\s*\)\s*,\s*\(\s*uppercase\s*,\s*r("|')\[A-Z\]\3\s*\)\s*,\s*\(\s*special_chars\s*,\s*r("|')\[\^a-zA-Z0-9\]\4\s*\)\s*,?\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = r'\.'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
