---
id: 65646ffeaed2d238c562a014
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

Es una convención común colocar las sentencias `import` al inicio de tu código. Y además, en caso de múltiples sentencias `import`, ordénalas en orden alfabético para mejorar la legibilidad.

En la parte superior de tu código, importa el módulo `random`.

# --hints--

Debes importar el módulo `random`.

```js
({ test: () => assert.match(code, /^import\s+random/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string
--fcc-editable-region--

# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)

```
