---
id: 65649a631440b757f9f7b428
title: Paso 28
challengeType: 20
dashedName: step-28
---

# --description--

Una expresión regular, o regex, es un patrón usado para coincidir con una combinación específica de caracteres dentro de una cadena. Es una alternativa válida a las sentencias condicionales `if`/`else` cuando necesitas coincidir o encontrar patrones dentro de una cadena para propósitos de validación, reemplazo de caracteres y otros.

La función `compile()` del módulo `re` compila la cadena pasada como argumento en un objeto de expresión regular que puede ser usado por otros métodos de `re`.

Declara una nueva variable `pattern` y asigna el valor de `re.compile('i')` a esta variable.

# --hints--

Debes declarar una variable `pattern`.

```js
({ test: () => assert(__userGlobals.has("pattern")) })
```

Debes asignar `re.compile('i')` a tu variable `pattern`.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('i')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--

--fcc-editable-region--
```
