---
id: 6564642ce403fe29cc23fee7
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

Un módulo de Python es un archivo que contiene código diseñado para realizar tareas específicas. La biblioteca estándar de Python contiene muchos módulos que puedes importar y usar en tu código. Puedes lograr esto usando la sentencia `import` seguida del nombre del módulo.

Comienza este proyecto importando el módulo `string`.

# --hints--

Debes importar el módulo `string`.

```js
({ test: () => assert.match(code, /^import\s+string/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
