---
id: 6564683821b2ee3174e7250e
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Puedes acceder a las utilidades definidas dentro del módulo importado mediante la notación de punto. La notación de punto funciona añadiendo un punto seguido del nombre de la utilidad al nombre del módulo. Por ejemplo, así es como se accede a la constante `ascii_lowercase`:

```py
import string


print(string.ascii_lowercase)
# Output: abcdefghijklmnopqrstuvwxyz
```

En este proyecto, vas a usar diferentes constantes del módulo `string`. 

Declara una nueva variable llamada `letters` y asigna `string.ascii_letters` a esta variable.

# --hints--

Debes declarar una variable llamada `letters`.

```js
({ test: () => assert(__userGlobals.has("letters")) })
```

Debes asignar `string.ascii_letters` a tu variable `letters`.

```js
({ test: () => assert(runPython(`
    import string
    letters == string.ascii_letters
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


--fcc-editable-region--
```
