---
id: 656485a8a3496d4a36b1496a
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Estas tres variables constituyen los caracteres posibles para elegir al generar la contraseña.

Justo antes de ellos, añade un comentario que diga `Define the possible characters for the password`.

# --hints--

Debes agregar el comentario justo encima de la variable `letters`.

```js
({ test: () => assert.match(code, /^#\s*Define the possible characters for the password.*^letters/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string



letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--
```
