---
id: 65648f4b2281ba50051ae39c
title: Paso 17
challengeType: 20
dashedName: step-17
---

# --description--

A continuación, escribe un bucle `for` con `i` como la variable del bucle. Usa la función `range()` para iterar hasta el valor de `length`.

Dentro del ciclo, usa el operador de asignación de suma para agregar un carácter aleatorio de `all_characters` al valor actual de `password`. Usa la función `choice()` del módulo `secrets` para eso.

# --hints--

Debes escribir un ciclo `for` que itere sobre `range(length)`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;    
    assert(function_body.match(/for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:/));
  }
})
```

Debes usar el operador `+=` para agregar un carácter aleatorio de `all_characters` al valor actual de `password`.

```js
({ test: () =>
  {
    const generate_pwd = __helpers.python.getDef(code, "generate_password");
    const {function_body} = generate_pwd;    
    assert(function_body.match(/^(\s*)for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:\s*^\1\1password\s*\+=\s*secrets\.choice\s*\(\s*all_characters\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--    
    password = ''
    # Generate password
    
--fcc-editable-region--
```
