---
id: 6564883669b5af4b69f794cc
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Tu variable `all_characters` es una cadena formada por todas las letras minúsculas y mayúsculas, los 10 dígitos y varios caracteres especiales. 

Justo antes, agrega un comentario que diga `Combine all characters`.

# --hints--

Debes agregar el comentario justo encima de tu variable `all_characters`.

```js
({ test: () => assert.match(code, /^#\s*Combine all characters\s+^all_characters/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation


all_characters = letters + digits + symbols
--fcc-editable-region--
```
