---
id: 65649122c7f77f519aaf0975
title: Paso 20
challengeType: 20
dashedName: step-20
---

# --description--

Finalmente, declara una variable `new_password` y asígnale el resultado de llamar a `generate_password`. Pasa `8` como argumento a tu llamada a `generate_password`.

# --hints--

Debes llamar a `generate_password` pasando `8` como argumento.

```js
({ test: () => assert.match(code, /generate_password\s*\(\s*8\s*\)/) })
```

Debes asignar `generate_password(8)` a la variable `new_password`.

```js
({
  test: () => {
    assert(__userGlobals.has("new_password"));
    assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*8\s*\)/m);
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--

--fcc-editable-region--
```
