---
id: 656ef5bc5c0cd464be1df675
title: Pasos 23
challengeType: 20
dashedName: step-23
---

# --description--

A continuación, vas a darle a tu función más parámetros que actuarán como restricciones para la contraseña generada.

Modifica la declaración de tu función agregando `nums`, `special_chars`, `uppercase` y `lowercase` en este orden después del parámetro `length` existente.

# --hints--

Tu función debe tomar `length`, `nums`, `special_chars`, `uppercase` y `lowercase` como parámetros. El orden importa.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(generate_password))
    sig == '(length, nums, special_chars, uppercase, lowercase)'
  `))
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

--fcc-editable-region--
def generate_password(length):
--fcc-editable-region--    
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password

# new_password = generate_password(8)
# print(new_password)
```
