---
id: 656ef89b4b486b7a16077864
title: Paso 26
challengeType: 20
dashedName: step-26
---

# --description--

Una tupla es otra estructura de datos incorporada en Python. Las tuplas son muy similares a las listas, pero se definen con paréntesis `()`, en lugar de corchetes. Además, las tuplas son inmutables, a diferencia de las listas.

```py
my_tuple = ('larch', 1, True)
```

Tu lista de `constraints` va a almacenar tuplas. El primer elemento de cada tupla será un parámetro de restricción.

Modifica la asignación de la lista `constraints` agregando una tupla a tu lista. Usa `nums` como el primer elemento y una cadena vacía como el segundo elemento.

# --hints--

Debes agregar una tupla a la lista `constraints`.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(.*\)\s*\]/) })
```

Debes agregar `(nums, '')` a la lista `constraints`.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*("|')\1\s*\)\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = []
        
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
