---
id: 656ef54070c72161e6feb90f
title: Paso 22
challengeType: 20
dashedName: step-22
---

# --description--

Todo parece bien, pero sería bueno tener una forma de comprobar que la contraseña generada cumple con características específicas. Por ejemplo, un número mínimo de caracteres especiales, dígitos, o letras mayúsculas/minúsculas. Muy pronto se ocuparán de ello.

Por ahora, comente las dos últimas líneas de su código.

# --hints--

Debes convertir las dos últimas líneas de tu código en comentarios.

```js
({ test: () => {
  assert.match(code, /#\s*new_password\s*=\s*generate_password\s*\(\s*8\s*\)/);
  assert.match(code, /#\s*print\s*\(\s*new_password\s*\)/);
} })
```


# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--
new_password = generate_password(8)
print(new_password)
--fcc-editable-region--
```
