---
id: 656f0d6bc60b58b9777a36c9
title: Paso 34
challengeType: 20
dashedName: step-34
---

# --description--

Una clase de caracteres se indica con corchetes y coincide con un carácter entre los especificados dentro de los corchetes. Por ejemplo, `ma[dnt]` coincide con `mad`, `man` o `mat`.

Modifica tu patrón para que coincida con una `w` seguida de `h` o `a`.

# --hints--

Tu patrón debe usar una clase de caracteres para coincidir con un `w` seguido de `h` o `a`.

```js
({ test: () => assert(runPython(`
    pattern == "w[ha]" or pattern == "w[ah]"    
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'l+'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
