---
id: 656f10458134d4c4e283a2f1
title: Paso 35
challengeType: 20
dashedName: step-35
---

# --description--

Ahora, convierte la cadena vacía en la tupla de restricción en un patrón regex para coincidir con un solo dígito. Usa una clase de caracteres que especifique cada dígito desde `0` hasta `9`.

# --hints--

Debes tener la tupla `(nums, '[0123456789]')` en tu lista de `constraints`.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\s*\(\s*nums\s*,\s*("|')\[\d{10}\]\1\s*\)\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, '')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = 'w[ha]'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
