---
id: 656f0bc259c1f6b57486ed68
title: Paso 33
challengeType: 20
dashedName: step-33
---

# --description--

Para verificar que la contraseña generada cumpla con las características requeridas, vas a usar la función `findall()` del módulo `re`. Es similar a `search` pero devuelve una lista con todas las ocurrencias del patrón encontrado.

Reemplaza la llamada a `search()` por `findall()` y verifica la salida.

# --hints--

Debes modificar tu llamada existente a `print()` reemplazando `search()` con `findall()`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'l+'
quote = 'Not all those who wander are lost.'
print(re.search(pattern, quote))
--fcc-editable-region--
```
