---
id: 6577319039f4f7de9251b822
title: Paso 61
challengeType: 20
dashedName: step-61
---

# --description--

En lugar de usar un ciclo y una variable contador, puedes lograr el mismo resultado con un enfoque diferente, que vas a implementar en los próximos pasos.

`all()` es una función incorporada de Python que devuelve `True` si todos los elementos dentro de un iterable dado evalúan a `True`. De lo contrario, devuelve `False`.

Reemplaza tu `for` loop existente y dos `if` statements con un solo `if` statement. Para la condición del `if`, usa una llamada a la función `all()` y pasa una lista vacía como argumento de la llamada a la función.

# --hints--

Debes reemplazar tu `for` loop y dos `if` statements existentes con un solo `if` statement. 

```js
({ test: () => {
  assert.match(code, /^(\s{8})if\s+.+:\s*^\1\s{4}break/ms);
  assert.isFalse(/if\s+count\s*==\s*4\s*:/.test(code));
} })
```

Tu nueva condición `if` debe ser `all([])`.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])") != _Node()
`)) })
```

Debes tener `break` dentro de la nueva estructura `if`.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])").find_bodies()[0].is_equivalent("break")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        for constraint, pattern in constraints:
            if constraint <= len(re.findall(pattern, password)):
                count += 1
            
        if count == 4:
--fcc-editable-region--
            break

    return password

# new_password = generate_password(8)
# print(new_password)
```
