---
id: 6577333feab1e8e927014f03
title: Paso 63
challengeType: 20
dashedName: step-63
---

# --description--

Teniendo `all([expression for i in iterable])`, significa que una nueva lista se crea evaluando `expressión` para cada `i` en `iterable`. Después de la función `all()` iterará sobre la lista recién creada, la lista se elimina automáticamente, ya que ya no es necesaria.

La memoria se puede guardar usando una expresión generadora. Las expresiones de generador siguen la sintaxis de las conversiones de listas, pero utilizan paréntesis en lugar de corchetes.

Cambie su lista de problemas en una expresión generadora eliminando los corchetes.

# --hints--

Debes convertir tu lista en una expresión generadora eliminando los corchetes.

```js
({ test: () => assert.match(code, /^(\s*)if\s+all\s*\(\s*constraint\s*<=\s*len\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*password\s*\)\s*\)\s*for\s+constraint\s*,\s*pattern\s+in\s+constraints\s*\)\s*:\s*^\1\s{4}break/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        if all(
            [
                constraint <= len(re.findall(pattern, password))
                for constraint, pattern in constraints
            ]
        ):
--fcc-editable-region--
            break
    
    return password

# new_password = generate_password(8)
# print(new_password)
```
