---
id: 657db4cb77190e33a20e852a
title: Paso 69
challengeType: 20
dashedName: step-69
---

# --description--

Cuando combinas argumentos por defecto con argumentos por palabra clave, puedes pasar explícitamente menos argumentos de los que requiere la función. Los argumentos que no se pasan explícitamente en la llamada a la función tomarán sus valores por defecto.

Modifica tu llamada a `generate_password()` para que tome solo `length=8`.

# --hints--

Debes pasar solo `length=8` a tu llamada `generate_password()`.

```js
({ test: () => assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*length\s*=\s*8\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length=16, nums=1, special_chars=1, uppercase=1, lowercase=1):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (special_chars, fr'[{symbols}]'),
            (uppercase, r'[A-Z]'),
            (lowercase, r'[a-z]')
        ]

        # Check constraints        
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):
            break
    
    return password
    
--fcc-editable-region--
new_password = generate_password(nums=1, length=8, special_chars=1, uppercase=1, lowercase=1)
print(new_password)
--fcc-editable-region--
```
