---
id: 67c562286b29447da020d407
title: Crea un Generador de Contraseña de Un Solo Uso (OTP)
challengeType: 25
dashedName: build-a-one-time-password-generator
demoType: onClick
---

# --description--

En este laboratorio, generarás un OTP (Contraseña de Un Solo Uso) de 6 dígitos y se lo mostrarás al usuario. El OTP expirará después de 5 segundos, y se generará uno nuevo cuando el usuario haga clic en el botón `"Generar nuevo OTP"`.

**Objetivo:** Cumple los relatos de usuario a continuación y pasa todas las pruebas para completar el laboratorio.

**Relatos de Usuario:**

1. Debes usar el gancho `useEffect` para gestionar el temporizador de cuenta regresiva.
2. Tu componente `OTPGenerator` debe devolver un elemento `div` con el nombre de clase `container`.
3. El `div` con la clase `container` debe incluir los siguientes elementos:

    - Un elemento `h1` con el ID `otp-title` y texto `"Generador de OTP"`.
    - Un elemento `h2` con el ID `otp-display` que muestra el mensaje `"Haz clic en 'Generar OTP' para obtener un código"` o muestra el OTP generado si está disponible.
    - Un elemento `p` con el ID `otp-timer` que:
        - Empieza vacío.
        - Muestra `"Expira en: X segundos"` después de hacer clic en el botón, donde `X` representa el tiempo restante antes de que el OTP expire.
        - Muestra el mensaje `"OTP expirado. Haz clic en el botón para generar un nuevo OTP."` una vez que la cuenta regresiva llegue a `0`.
    - Un elemento `button` con el ID `generate-otp-button` etiquetado `"Generar OTP"`. Al hacer clic, debe generar un nuevo OTP y comenzar una cuenta regresiva de 5 segundos.
    - El botón `"Generar OTP"` debe estar deshabilitado mientras la cuenta regresiva esté activa.

4. Debe garantizar que el temporizador de cuenta regresiva se detenga automáticamente una vez que llegue a `0` segundos para evitar actualizaciones innecesarias.
5. El OTP generado debe tener 6 dígitos.

# --before-all--

```js
let clock = __FakeTimers.install();
```

# --hints--

No debes eliminar la asignación existente `const { useState, useEffect, useRef } = React;` del código.

```js
assert.match(code, /{\s*(useState\s*,\s*useEffect\s*,\s*useRef|useState\s*,\s*useRef\s*,\s*useEffect|useEffect\s*,\s*useState\s*,\s*useRef|useEffect\s*,\s*useRef\s*,\s*useState|useRef\s*,\s*useState\s*,\s*useEffect|useRef\s*,\s*useEffect\s*,\s*useState|useState|useEffect|useRef)\s*}\s*=\s*React;/);
```

Debes renderizar un `div` con la clase `.container`.

```js
const containerDiv = document.querySelector('.container');
assert.exists(containerDiv);
```

El `.container` debe contener un elemento `h1` con el ID `otp-title` y el texto `"Generador de OTP"`.

```js
const h1 = document.querySelector(".container h1#otp-title");
assert.exists(h1);
assert.equal(h1.textContent, "OTP Generator");
```

El `.container` debe contener un elemento `h2` con el ID `otp-display` que muestra el OTP cuando se genera.

```js
const h2 = document.querySelector(".container h2#otp-display");
assert.exists(h2);
```

Inicialmente, el `h2` dentro de `.container` debe mostrar el mensaje `"Haz clic en 'Generar OTP' para obtener un código"`.

```js
assert.strictEqual(document.querySelector('.container h2#otp-display').textContent.trim(), "Click 'Generate OTP' to get a code");
```

Inicialmente, debe haber un elemento `.container` que no contenga ningún elemento `<p>`.

```js
const container = document.querySelector('.container');
assert.exists(container, "The container element should exist");
const pElements = document.querySelectorAll('.container p');
assert.strictEqual(pElements.length, 0, "Initially, there should be no <p> elements inside '.container'");
```

El `.container` debe contener un elemento `button` con el ID `generate-otp-button` y texto `"Generar OTP"`.

```js
const button = document.querySelector('.container button#generate-otp-button').textContent;
assert.exists(button);
assert.strictEqual(button, 'Generate OTP');
```

Cuando se haga clic en el botón, debería mostrar un nuevo OTP en el elemento `h2` con id `otp-display`.

```js
async() => {
try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)

assert.match(h2.textContent?.trim(), /[0-9]/, "OTP should be a number");
} finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
}
```

El OTP generado debe tener 6 dígitos.

```js
async() => {
  try {
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1);
const h2 = document.querySelector(".container h2#otp-display");

assert.match(h2.textContent.trim(), /^[0-9]{6}$/, "OTP should be a 6-digit number");

 } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
}
```


Cuando se haga clic en el botón, el elemento `p` con id de `otp-timer` debe mostrar una cuenta regresiva de 5 segundos.

```js
async() => {
  try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)
const p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/, "Countdown should start at 5 seconds");

  } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
}
```

El mensaje en el elemento `p` con id `otp-timer` debe actualizarse cada segundo para mostrar el tiempo restante.

```js
async() => {
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1)
let p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/);
await clock.tickAsync(1000)
p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 4 seconds/, "Countdown should update to 4 seconds after 1 second");
}
```

El botón `"Generar OTP"` debe estar deshabilitado mientras el temporizador de cuenta regresiva está activo.

```js
async () => {
try {
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();
  clock.tick(1);

  assert.isTrue(button.disabled);
} finally {
  await clock.runAllAsync();
}
}
```

El botón `"Generar OTP"` debe estar habilitado una vez que el temporizador de cuenta regresiva llegue a 0.

```js
async () => {

const button = document.querySelector(".container button#generate-otp-button");
button.click();
await clock.tickAsync(6000);

assert.isFalse(button.disabled);
}
```

Cuando el temporizador de cuenta regresiva llegue a `0`, deberías mostrar el mensaje `OTP expirado. Haz clic en el botón para generar un nuevo OTP.`.

```js
async () => {
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();

  await clock.tickAsync(6000);

  const p = document.querySelector(".container p#otp-timer");
  assert.equal(p.textContent.trim(), "OTP expired. Click the button to generate a new OTP.");
}
```

Debes exportar el componente `OTPGenerator`.

```js
assert.isFunction(window.index.OTPGenerator);
```

Desinstala el reloj, ignora esta prueba

```js
async () => {
clock.uninstall()
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Color Picker</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css

```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {};
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css
body {
    font-family: 'Castoro Titling', cursive;
}

.container {
    text-align: center;
    margin-top: 50px;
}

h1 {
    font-size: 30px;
    color: #333;
}

h2 {
    font-size: 25px;
    color: #555;
}

p {
    font-size: 16px;
    color: #888;
}

button {
    padding: 10px 20px;
    margin-top: 20px;
    font-size: 16px;
    border: none;
    background-color: #383f59;
    color: white;
    border-radius: 5px;
}

button:disabled {
  background-color:#57585c;
  cursor: not-allowed;
}

button:hover {
    background-color: #444;
}
```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {
  const intervalId = useRef(null);
  const [otp, setOtp] = useState(null);
  const [secondsLeft, setSecondsLeft] = useState(5);
  const [isCounting, setIsCounting] = useState(false);
  const [hasGeneratedOtp, setHasGeneratedOtp] = useState(false);

  const generateOTP = () => Math.floor(100000 + Math.random() * 900000);

  const handleGenerateOtp = () => {
    setOtp(generateOTP());
    setSecondsLeft(5);
    setIsCounting(true);
    setHasGeneratedOtp(true);
    intervalId.current = setInterval(() => {
      setSecondsLeft(prevSeconds => prevSeconds - 1);
    }, 1000);
  };

  useEffect(() => {
    if (secondsLeft === 0) {
      setIsCounting(false);
      clearInterval(intervalId.current);
    }
  }, [secondsLeft]);

  useEffect(
    () => () => {
      clearInterval(intervalId.current);
    },
    []
  );

  return (
    <div className='container'>
      <h1 id='otp-title'>OTP Generator</h1>
      <h2 id='otp-display'>
        {otp ? otp : "Click 'Generate OTP' to get a code"}
      </h2>
      {isCounting ? (
        <p id='otp-timer'>Expires in: {secondsLeft} seconds</p>
      ) : (
        hasGeneratedOtp && (
          <p id='otp-timer'>
            OTP expired. Click the button to generate a new OTP.
          </p>
        )
      )}
      <button
        id='generate-otp-button'
        onClick={handleGenerateOtp}
        disabled={isCounting}
      >
        Generate OTP
      </button>
    </div>
  );
};
```
