---
id: 60a3e3396c7b40068ad69971
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

CSSの時間です。

`<div>` にテキストがなくても、内容のあるボックスとして扱われます。
`.canvas` クラスセレクターを使い、その `width` を500ピクセルに設定するCSSルールを書いてください。
以下は、クラス `card` の幅を300ピクセルに設定するCSSルールの例です。

```css
.card {
  width: 300px;
}
```

# --hints--

コードには `.canvas` セレクターが必要です。

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

`width` プロパティを `500px` に設定してください。

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

`.canvas` 要素の `width` は `500px` であるべきです。

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
