---
id: 60a3e3396c7b40068ad69975
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

`.frame` クラスセレクターを使って新しいルールを書いてください。

`border` ショートハンド宣言を使って、`.frame` 要素に幅 `50px` の黒の実線ボーダーを設定してください。

# --hints--

コードには `.frame` セレクターが必要です。

```js
const hasFrame = new __helpers.CSSHelp(document).getStyle('.frame');
assert(hasFrame);
```

`border` プロパティを `50px solid black` に設定してください。

```js
const hasBorder = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.border === '50px solid black');
assert(hasBorder);
```

`.frame` 要素には `50px solid black` `border` が必要です。

```js
const frameBorder = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('border');
assert(frameBorder === '50px solid black');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
--fcc-editable-region--

--fcc-editable-region--
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
