---
id: 60a3e3396c7b40068ad6996a
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

ここまでで、HTMLページに必要な基本的な要素に慣れているはずです。

`DOCTYPE` 宣言、言語が英語に設定された `html` 要素、`head` 要素、そして `body` 要素を使ってコードを設定してください。

# --hints--

コードには `<!DOCTYPE html>` 宣言が必要です。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

開く `<html>` タグには `lang` 属性が `en` に設定されている必要があります。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

コードには `head` 要素が `html` 要素の中にある必要があります。

```js
assert(document.querySelectorAll('head').length === 1);
```

コードには `body` 要素が `html` 要素の中にある必要があります。

```js
assert(document.querySelectorAll('body').length === 1);
```

`head` 要素は `body` 要素の前に配置してください。

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
