---
id: 60a3e3396c7b40068ad6998b
title: ステップ 33
challengeType: 0
dashedName: step-33
---

# --description--

マージンを一方向に押し込むのは便利です。

この場合、`.one` 要素の下マージンが `.two` を20ピクセル下に押し下げています。

`.two` セレクターでは、`margin` 省略形プロパティを使って、上マージンを `0`、左右のマージンを `auto`、下マージンを `20px` に設定してください。これにより、上マージンがなくなり、左右中央に配置され、下マージンが20ピクセルに設定されます。

# --hints--

`margin` プロパティを `0 auto 20px` に設定してください。

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '0px auto 20px');
assert(hasMargin);
```

`.two` 要素は `margin` の値を `0 auto 20px` にする必要があります。

```js
const twoMargin = new __helpers.CSSHelp(document).getStyle('.two')?.getPropertyValue('margin');
assert(twoMargin === '0px auto 20px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
