---
id: 60a3e3396c7b40068ad6996d
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

コンテンツはパディングと呼ばれるスペースで囲まれており、これはプチプチが中身と箱の間にあるのと似ています。

ボーダーは、商品が配送された段ボール箱のようなものだと考えてください。

`src` 属性を `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png` に変更してください。

# --hints--

`img` 要素は `src` の `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png` を持つべきです。

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```
