---
id: 60a3e3396c7b40068ad6996b
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`head` 要素の中に、`meta` を `charset` に設定する `UTF-8` タグと、値が `title` の `Rothko Painting` 要素を追加してください。

`body` 要素の中に、`img` の `src` を持つ `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` 要素を追加してください。
# --hints--

コードには `meta` タグが含まれている必要があります。

```js
assert(document.querySelectorAll('meta').length === 1);
```

`meta` タグは `charset` を `UTF-8` に設定する必要があります。

```js
assert(document.querySelector('meta').getAttribute('charset')?.toLowerCase() === 'utf-8');
```

コードには `title` 要素が含まれている必要があります。

```js
assert(document.querySelectorAll('title').length === 1);
```

`title` は `Rothko Painting` である必要があります。

```js
assert(document.querySelector('title').innerText === 'Rothko Painting');
```

コードには `img` 要素が含まれている必要があります。

```js
assert(document.querySelectorAll('img').length === 1);
```

`img` 要素は `src` の `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` を持つ必要があります。

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
