---
id: 60a3e3396c7b40068ad6996c
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

CSSボックスモデルでは、すべてのHTML要素が4つのエリアを持つボックスとして扱われます。

お気に入りのオンライン小売店から箱を受け取ることを想像してください。箱の中身は商品であり、ここではヘッダー、段落、または画像要素にあたります。

`src`の`<img>`属性を`https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`から`https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`に変更してください。

# --hints--

`img`要素は`src`の`https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`を持つべきです

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```
