---
id: 61fd67a656743144844941cb
title: Hatua ya 4
challengeType: 0
dashedName: step-4
---

# --description--

Screen readers announce HTML elements based on the document flow. We will eventually want the balance sheet to have a heading of "Balance Sheet" and a subheading of "AcmeWidgetCorp". However, this order does not make sense if announced by a screen reader.

Ipe `span` yako iliyopo sifa ya `class` iliyowekwa kuwa `flex`, na uongeze vipengele viwili vya `span` ndani yake. Ipe ya kwanza maandishi `AcmeWidgetCorp`. Ipe ya pili maandishi `Balance Sheet`. Utatumia CSS kubadilisha mpangilio wa maandishi kwenye ukurasa, lakini mpangilio wa HTML utaleta maana zaidi kwa kisomaji skrini.

# --hints--

Kipengele chako cha `span` kilichopo kinafaa kuwa na sifa ya `class` iliyowekwa kuwa `flex`.

```js
assert.isTrue(
  document.querySelector('h1')?.children?.[0]?.classList?.contains('flex')
);
```

Kipengele chako cha `span` kilichopo kinafaa kuwa na vipengele viwili vipya vya `span` ndani yake.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.localName,
  'span'
);
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.localName,
  'span'
);
```

Vipengele vyako vipya vya `span` havifai kuwa na sifa ya `class`.

```js
assert.isEmpty(document.querySelector('.flex')?.children?.[0]?.classList);
assert.isEmpty(document.querySelector('.flex')?.children?.[1]?.classList);
```

Kipengele chako kipya cha kwanza cha `span` kinapaswa kuwa na maandishi `AcmeWidgetCorp`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.textContent,
  'AcmeWidgetCorp'
);
```

Kipengele chako cha pili cha `span` kinapaswa kuwa na maandishi `Balance Sheet`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.textContent,
  'Balance Sheet'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
