---
id: 61fd5a93fd62bb35968adeab
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Anza mradi wako kwa kumpa kipengele chako cha `head` kipengele cha `link` kwa ajili ya karatasi yako ya mitindo.

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `link`.

```js
assert.match(code, /<link/);
```

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele chako cha `head`.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele ya `rel` yenye thamani `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/);
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele ya `href` yenye thamani `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
