---
id: 61fd78621573aa5e8b512f5e
title: Hatua 15
challengeType: 0
dashedName: step-15
---

# --description--

Katika kipengele chako cha tatu cha `tr`, ongeza kipengele cha `th` chenye maandishi ya `Savings Funds set aside for emergencies.`. Zungusha maandishi hayo, isipokuwa `Savings `, katika kipengele cha `span` chenye sifa ya `class` iliyowekwa kuwa `description`.

Baada ya hapo, ongeza vipengele vitatu vya `td` vyenye maandishi yafuatayo (kwa mpangilio): `$500`, `$650`, `$728`. Toa kipengele cha tatu cha `td` sifa ya `class` iliyowekwa kuwa `current`.

# --hints--

Kipengele chako cha tatu cha `tr` kinapaswa kuwa na kipengele cha `th`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.isNotNull(tableRow?.querySelector('th'));
```

Kipengele chako cha `th` kinapaswa kuwa na maandishi ya `Savings Funds set aside for emergencies.`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(
  tableRow?.querySelector('th')?.innerText.trim(),
  'Savings Funds set aside for emergencies.'
);
```

Unapaswa kuzungusha maandishi ya `Funds set aside for emergencies.` katika kipengele cha `span`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(
  tableRow?.querySelector('th > span')?.textContent.trim(),
  'Funds set aside for emergencies.'
);
```

Kipengele chako cha `span` kinapaswa kuwa na sifa ya kipengele ya `class` iliyowekwa kuwa `description`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.isTrue(
  tableRow?.querySelector('th > span')?.classList?.contains('description')
);
```

Unapaswa kuwa na vipengele vitatu vya `td`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.lengthOf(tableRow?.querySelectorAll('td'), 3);
```

Kipengele chako cha `td` cha kwanza kinapaswa kuwa na maandishi `$500`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[0]?.textContent.trim(), '$500');
```

Kipengele chako cha pili cha `td` kinapaswa kuwa na maandishi ya `$650`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[1]?.textContent.trim(), '$650');
```

Kipengele chako cha tatu cha `td` kinapaswa kuwa na maandishi ya `$728`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[2]?.textContent.trim(), '$728');
```

Kipengele chako cha tatu cha `td` kinapaswa kuwa na `class` iliyowekwa kuwa `current`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.isTrue(
  tableRow?.querySelectorAll('td')?.[2]?.classList?.contains('current')
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
--fcc-editable-region--
                
--fcc-editable-region--
              </tr>
              <tr class="total"></tr>
            </tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
