---
id: 61fd71d596e8f253b9408b39
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Kipengele cha `tr` kinatumika kuonyesha safu ya jedwali. Ongeza kipengele cha `tr` ndani ya kipengele chako cha `thead`. Katika kipengele chako kipya cha `tr`, ongeza kipengele cha `td`, kifuatiwe na vipengele vitatu vya `th`.

Kipengele cha `td` kinaonyesha seli ya data, wakati kipengele cha `th` kinaonyesha seli ya kichwa.

# --hints--

Kipengele chako cha `thead` kinapaswa kuwa na kipengele cha `tr`.

```js
assert.strictEqual(
  document.querySelector('thead')?.children?.[0]?.localName,
  'tr'
);
```

Kipengele chako cha `tr` kinapaswa kuwa na kipengele cha `td` kama mtoto wa kwanza.

```js
assert.strictEqual(
  document.querySelector('tr')?.children?.[0]?.localName,
  'td'
);
```

Kipengele chako cha `tr` kinapaswa kuwa na vipengele vitatu vya `th`, baada ya kipengele cha `td`.

```js
assert.strictEqual(
  document.querySelector('tr')?.children?.[1]?.localName,
  'th'
);
assert.strictEqual(
  document.querySelector('tr')?.children?.[2]?.localName,
  'th'
);
assert.strictEqual(
  document.querySelector('tr')?.children?.[3]?.localName,
  'th'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
            </thead>
--fcc-editable-region--
            <tbody></tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
