---
id: 61fd9ad665a4a282c8106be3
title: Hatua 26
challengeType: 0
dashedName: step-26
---

# --description--

Ndani ya `thead`, tengeneza kipengele cha `tr`. Ndani yake, ongeza `td` na vipengele vitatu vya `th`. Ndani ya kila kipengele cha `th`, ongeza kipengele cha `span` chenye sifa ya kipengele `class` iliyowekwa kuwa `sr-only` na maandishi yafuatayo, kwa mpangilio: `2019`, `2020`, na `2021`.

# --hints--

Kipengele chako cha `thead` kinapaswa kuwa na kipengele cha `tr`.

```js
const table = document.querySelectorAll('table')?.[2];
assert.strictEqual(
  table?.querySelector('thead')?.children?.[0]?.localName,
  'tr'
);
```

Kipengele chako cha `tr` kinapaswa kuwa na kipengele cha `td` kama mtoto wa kwanza.

```js
const table = document.querySelectorAll('table')?.[2];
assert.strictEqual(table?.querySelector('tr')?.children?.[0]?.localName, 'td');
```

Kipengele chako cha `tr` kinapaswa kuwa na vipengele vitatu vya `th`, baada ya kipengele cha `td`.

```js
const table = document.querySelectorAll('table')?.[2];
const trElement = table?.querySelector('tr');
assert.strictEqual(trElement?.children?.[1]?.localName, 'th');
assert.strictEqual(trElement?.children?.[2]?.localName, 'th');
assert.strictEqual(trElement?.children?.[3]?.localName, 'th');
```

Kila kipengele chako cha `th` kinapaswa kuwa na kipengele cha `span`.

```js
const ths = [
  ...document.querySelectorAll('table')?.[2]?.querySelectorAll('th')
];
ths?.forEach(th => {
  assert.lengthOf(th?.children, 1);
  assert.strictEqual(th?.children?.[0]?.localName, 'span');
});
```

Kila kipengele chako kipya cha `span` kinapaswa kuwa na sifa ya kipengele `class` iliyowekwa kuwa `sr-only`.

```js
const ths = [
  ...document.querySelectorAll('table')?.[2]?.querySelectorAll('th')
];
ths?.forEach(th => {
  assert.isTrue(th?.children?.[0]?.classList?.contains('sr-only'));
});
```

Kipengele chako cha kwanza cha `span` kinapaswa kuwa na maandishi `2019`.

```js
const table = document.querySelectorAll('table')?.[2];
const thead = table.querySelectorAll('th')?.[0];
assert.strictEqual(thead?.children?.[0]?.textContent.trim(), '2019');
```

Kipengele chako cha pili cha `span` kinapaswa kuwa na maandishi `2020`.

```js
const table = document.querySelectorAll('table')?.[2];
const thead = table.querySelectorAll('th')?.[1];
assert.strictEqual(thead?.children?.[0]?.textContent.trim(), '2020');
```

Kipengele chako cha tatu cha `span` kinapaswa kuwa na maandishi `2021`.

```js
const table = document.querySelectorAll('table')?.[2];
const thead = table.querySelectorAll('th')?.[2];
assert.strictEqual(thead?.children?.[0]?.textContent.trim(), '2021');
```

Kipengele chako cha `td` kinapaswa kuwa tupu.

```js
const table = document.querySelectorAll('table')?.[2];
const tableData = table?.querySelectorAll('td')?.[0];
assert.strictEqual(tableData?.textContent.trim(), '');
assert.isTrue(tableData?.children?.length === 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>Savings <span class="description">Funds set aside for emergencies.</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
              <tr class="total">
                <th>Total <span class="sr-only">Assets</span></th>
                <td>$579</td>
                <td>$736</td>
                <td class="current">$809</td>
              </tr>
            </tbody>
          </table>
          <table>
            <caption>
              Liabilities
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only">2019</span></th>
                <th><span class="sr-only">2020</span></th>
                <th><span class="sr-only">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Loans <span class="description">The outstanding balance on our startup loan.</span></th>
                <td>$500</td>
                <td>$250</td>
                <td class="current">$0</td>
              </tr>
              <tr class="data">
                <th>Expenses <span class="description">Annual anticipated expenses, such as payroll.</span></th>
                <td>$200</td>
                <td>$300</td>
                <td class="current">$400</td>
              </tr>
              <tr class="data">
                <th>Credit <span class="description">The outstanding balance on our credit card.</span></th>
                <td>$50</td>
                <td>$50</td>
                <td class="current">$75</td>
              </tr>
              <tr class="total">
                <th>Total <span class="sr-only">Liabilities</span></th>
                <td>$750</td>
                <td>$600</td>
                <td class="current">$475</td>
              </tr>
            </tbody>
          </table>
          <table>
            <caption>Net Worth</caption>
            <thead>
--fcc-editable-region--
              
--fcc-editable-region--
            </thead>
            <tbody>
            </tbody>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
