---
id: 61fd6b7c83dbf54a08cf0498
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Ndani ya kipengele chako cha `div`, ongeza vipengele vitatu vya `span`. Kila kimoja kitoe sifa ya kipengele ya `class` iliyowekwa kuwa `year`, na ongeza maandishi yafuatayo (kwa mpangilio): `2019`, `2020`, na `2021`.

# --hints--

Kipengele chako cha `div` kinapaswa kuwa na vipengele vitatu vya `span`.

```js
assert.lengthOf(document.querySelector('div')?.children, 3);
```

Kila kipengele cha `span` kinapaswa kuwa na sifa ya kipengele ya `class` iliyowekwa kuwa `year`.

```js
const spans = [...document.querySelector('div')?.children];
spans.forEach(span => assert.isTrue(span?.classList?.contains('year')));
```

`span` yako ya kwanza inapaswa kuwa na maandishi ya `2019`.

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[0]?.textContent.trim(),
  '2019'
);
```

`span` yako ya pili inapaswa kuwa na maandishi ya `2020`.

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[1]?.textContent.trim(),
  '2020'
);
```

`span` yako ya tatu inapaswa kuwa na maandishi ya `2021`.

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[2]?.textContent.trim(),
  '2021'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
--fcc-editable-region--
          
--fcc-editable-region--
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
