---
id: 61fd67a656743144844941cb
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Visomaji vya skrini hutangaza vipengele vya HTML kulingana na mtiririko wa hati. Hatimaye tutataka hesabu ya mizani iwe na kichwa cha sehemu cha "Balance Sheet" na kichwa kidogo cha "AcmeWidgetCorp". Hata hivyo, mpangilio huu hauwezi kueleweka vizuri ikiwa utatangazwa na kisomaji cha skrini.

Toa sifa ya `span` yenye thamani ya `class` kwa `flex` uliyonayo, kisha ongeza vipengele viwili vya `span` ndani yake. Mpe kipengele cha kwanza maandishi ya `AcmeWidgetCorp`. Mpe kipengele cha pili maandishi ya `Balance Sheet`. Utatumia CSS kubadilisha mpangilio wa maandishi kwenye ukurasa, lakini mpangilio wa HTML utakuwa na maana zaidi kwa kisomaji cha skrini.

# --hints--

Kipengele chako cha `span` kilichopo kinapaswa kuwa na sifa ya `class` yenye thamani ya `flex`.

```js
assert.isTrue(
  document.querySelector('h1')?.children?.[0]?.classList?.contains('flex')
);
```

Kipengele chako cha `span` kilichopo kinapaswa kuwa na vipengele viwili vipya vya `span` ndani yake.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.localName,
  'span'
);
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.localName,
  'span'
);
```

Vipengele vyako vipya vya `span` havipaswi kuwa na sifa ya `class`.

```js
assert.isEmpty(document.querySelector('.flex')?.children?.[0]?.classList);
assert.isEmpty(document.querySelector('.flex')?.children?.[1]?.classList);
```

Kipengele chako kipya cha kwanza cha `span` kinapaswa kuwa na maandishi ya `AcmeWidgetCorp`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.textContent.trim(),
  'AcmeWidgetCorp'
);
```

Kipengele chako kipya cha pili cha `span` kinapaswa kuwa na maandishi ya `Balance Sheet`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.textContent.trim(),
  'Balance Sheet'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
