---
id: 61fd7648a7ba2e5882436831
title: Hatua 12
challengeType: 0
dashedName: step-12
---

# --description--

Ndani ya kipengele chako cha `tbody`, ongeza vipengele vinne vya `tr`. Toa sifa ya kipengele ya `class` kwa vipengele vitatu vya kwanza na uiweke kuwa `data`, na sifa ya kipengele ya `class` kwa kipengele cha nne na uiweke kuwa `total`.

# --hints--

Kipengele chako cha `tbody` kinapaswa kuwa na vipengele vinne vya `tr`.

```js
const children = [...document.querySelector('tbody')?.children];
assert.lengthOf(children, 4);
children.forEach(child => assert.strictEqual(child?.localName, 'tr'));
```

Vipengele vitatu vya kwanza vya `tr` vinapaswa kuwa na sifa ya kipengele ya `class` iliyowekwa kuwa `data`.

```js
const children = [...document.querySelector('tbody')?.children];
children.forEach((child, index) => {
  if (index < 3) {
    assert.isTrue(child?.classList?.contains('data'));
  }
});
```

Kipengele cha nne cha `tr` kinapaswa kuwa na sifa ya kipengele ya `class` iliyowekwa kuwa `total`.

```js
const children = [...document.querySelector('tbody')?.children];
assert.isTrue(children?.[3]?.classList?.contains('total'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
--fcc-editable-region--
              
--fcc-editable-region--
            </tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
