---
id: 61fd6343fb42013d99bcd7f3
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Ndani ya kipengele chako cha `body`, panga ndani kipengele cha `section` ndani ya kipengele cha `main`.

# --hints--

Kipengele cha kwanza ndani ya kipengele cha `body` kinapaswa kuwa kipengele chako cha `main`.

```js
assert.strictEqual(
  document.querySelector('body')?.children?.[0]?.localName,
  'main'
);
```

Unapaswa kuwa na kipengele cha `section` ndani ya kipengele chako cha `main`.

```js
assert.strictEqual(
  document.querySelector('main')?.children?.[0]?.localName,
  'section'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css

```
