---
id: 61fd6cc9475a784b7776233e
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Chini ya kipengele chako cha `div`, ongeza kipengele kipya cha `div` chenye `class` iliyowekwa kuwa `table-wrap`. Hii itakuwa kontena la meza zako.

Ndani yake, ongeza vipengele vitatu vya `table`. Utatumia CSS kuvitengeneza kuwa meza moja, lakini kutumia meza tofauti kutasaidia kisomaji cha skrini kuelewa mtiririko wa hati.

# --hints--

Unapaswa kuunda kipengele kipya cha `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

Kipengele chako kipya cha `div` kinapaswa kuwa na `class` iliyowekwa kuwa `table-wrap`.

```js
assert.strictEqual(document.querySelector('.table-wrap')?.localName, 'div');
```

Kipengele chako cha `.table-wrap` kinapaswa kuja baada ya `div` uliyonayo.

```js
assert.isTrue(
  document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap')
);
```

Kipengele chako cha `.table-wrap` kinapaswa kuwa na vipengele vitatu vya `table`.

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert.lengthOf(children, 3);
children.forEach(child => assert.strictEqual(child?.localName, 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
