---
id: 612e89d254fe5d3df7d6693d
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Ahora que ha restablecido el modelo de caja HTML, también necesita aplicarlo a los elementos dentro de él. Para hacerlo, deberá establecer la propiedad `box-sizing` de todos los demás elementos en heredado, lo que garantiza que los elementos seleccionados adopten el mismo valor que su elemento padre.

También necesitará apuntar a los pseudo-elementos, que son palabras clave especiales que siguen a un selector. Los dos pseudo-elementos que usará son los pseudo-elementos `::before` y `::after`.

El selector `::before` crea un pseudo-elemento que es el primer hijo del elemento seleccionado, mientras que el selector `::after` crea un pseudo-elemento que es el último hijo del elemento seleccionado. Estos pseudo-elementos se utilizan a menudo para crear contenido cosmético, lo cual verá más adelante en este proyecto.

Por ahora, cree un selector CSS para apuntar a todos los elementos usando `*`, e incluya los pseudo-elementos `::before` y `::after`. Establezca la propiedad `box-sizing` a `inherit`.

# --hints--

Debería tener un selector `html`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

Su selector `html` debería tener la propiedad `box-sizing` establecida en `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

Debería tener un selector `*, ::before, ::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

Su selector `*, ::before, ::after` debería tener la propiedad `box-sizing` establecida en `inherit`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing, 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
