---
id: 612e89562043183c86df287c
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Los navegadores suelen aplicar valores de márgenes y relleno predeterminados a elementos específicos. Para asegurarse de que su diseño de piano se muestre correctamente, necesita restablecer el modelo de caja.

Agregue un selector de regla `html` a su archivo CSS y establezca la propiedad `box-sizing` en `border-box`.

# --hints--

Debe tener un selector `html`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

El selector `html` debe tener la propiedad `box-sizing` establecida en `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
