---
id: 612ea4c4993aba52ab4aa32e
title: Paso 16
challengeType: 0
dashedName: step-16
---

# --description--

Ahora es momento de usar los pseudo-selectores que configuraste anteriormente. Para crear las teclas negras, agrega un nuevo selector `.key.black--key::after`. Esto apunta a los elementos con las clases `key` y `black--key` y selecciona el pseudo-elemento creado después de estos elementos en el HTML.

En el nuevo selector, coloca la propiedad `background-color` a `#1d1e22`. Establece también la propiedad `content` en `""`. Esto hará que los pseudo-elementos queden vacíos.

La propiedad `content` se usa para definir o sobrescribir el contenido de un elemento. Por defecto, los pseudo-elementos creados por los selectores `::before` y `::after` están vacíos, lo que significa que no se representan en la página. Al establecer la propiedad `content` en una cadena vacía `""`, aseguras que los pseudo-elementos se representen, aunque sigan apareciendo vacíos.

Si deseas experimentar, intenta eliminar la propiedad `background-color` y establecer diferentes valores para la propiedad `content`, como `"♥"`. Recuerda deshacer estos cambios cuando termines, para que las pruebas sean aprobadas.

# --hints--

Debe tener un selector `.key.black--key::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

Su selector `.key.black--key::after` debe tener una propiedad `background-color` establecida en `#1d1e22`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor, 'rgb(29, 30, 34)');
```

Su selector `.key.black--key::after` debe tener una propiedad `content` establecida en `""`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content, '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
