---
id: 612e813b3ba67633222cbe54
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Recuerde, un atributo `class` puede contener múltiples valores. Para diferenciar entre tus teclas blancas y negras, agrega un segundo valor de `class` de `black--key` a tus segundo, tercero, quinto, sexto y séptimo elementos `.key`.

# --hints--

Tu segundo elemento `.key` también debe tener un `class` de `black--key`.

```js
const key = document.querySelectorAll('.key')[1];
assert.isTrue(key?.className.includes('black--key'));
```

Tu tercer `.key` debe tener `black--key` en el `class`.

```js
const third = document.querySelectorAll('.key')?.[2];
assert.isTrue(third?.classList?.contains('black--key'));
```

Tu quinto `.key` debe tener `black--key` en el `class`.

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert.isTrue(fifth?.classList?.contains('black--key'));
```

Tu sexto `.key` debe tener `black--key` en el `class`.

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert.isTrue(sixth?.classList.contains('black--key'));
```

Tu séptimo `.key` debe tener `black--key` en el `class`.

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert.isTrue(seventh?.classList?.contains('black--key'));
```

Debes tener cinco elementos `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 5);
```

Debes tener siete elementos `.key`.

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
