---
id: 612e7d1c29fb872d6384379c
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Anide un segundo `div` dentro de su `div` existente y establezca la `class` en `keys`.

# --hints--

Debes crear un segundo elemento `div`.

```js
const divDiv = document.querySelectorAll('div');
assert.lengthOf(divDiv, 2);
```

Su nuevo elemento `div` debe estar dentro de su elemento `div` existente.

```js
const div = document.querySelector('div');
assert.lengthOf(div?.children, 1);
assert.equal(div?.children?.[0]?.localName, 'div');
```

Su nuevo elemento `div` debe tener la `class` establecida en `keys`.

```js
const div = document.querySelector('div');
assert.equal(div?.children?.[0]?.className, 'keys');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
  --fcc-editable-region--

  --fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
