---
id: 685239f0cad43c7a909cbd98
title: 调试宠物领养页面
challengeType: 25
dashedName: lab-debug-pet-adoption-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

Sally 是一位宠物领养商店的老板，她搭建了自己的第一个网页，但存在一些问题。

你的任务是修复所有错误，以便 Sally 能够继续构建她的页面。 

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. Sally 想使用一张猫咪的图像，但图像显示不正确。 你需要在 `img` 元素中修正以下内容：
   - 用图像源的正确属性替换 `href` 属性。
   - 将 `att` 属性替换为代表图像简短描述性文本的正确属性。
   - 删除 `</img>` 结束标签，因为 `img` 元素是空元素且没有结束标签。
2. Sally 希望使用一些链接将用户引导到猫和狗的页面。 但链接无法正常工作。 你需要在 `a` 元素中修正以下内容：
   - 将两个 `src` 属性替换为用于指定 URL 的正确属性。

# --hints--

你的 `img` 元素应该有一个 `src` 属性，而不是 `href` 属性。

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("src"));
```

你的 `img` 元素应该有一个 `alt` 属性，而不是不存在的 `att` 属性。

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("alt"));
```

你的 `img` 元素不应有 `</img>` 关闭标签。

```js
assert.notMatch(code, /<\/img>/);
```

你的文本为 `Visit cats page` 的 `a` 元素应该有一个 `href` 属性，而不是 `src` 属性。

```js
const anchors = document.querySelectorAll("a");
const catLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit cats page");
assert.isDefined(catLink);
assert.isTrue(catLink.hasAttribute("href"));
```

你的文本为 `Visit dogs page` 的 `a` 元素应该有一个 `href` 属性，而不是 `src` 属性。

```js
const anchors = document.querySelectorAll("a");
const dogLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit dogs page");
assert.isDefined(dogLink);
assert.isTrue(dogLink.hasAttribute("href"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" att="Two tabby kittens sleeping together on a couch."></img>

<h2>Adopt a cat!</h2>
<a src="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a src="/dogs">Visit dogs page</a>
```

# --solutions--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">

<h2>Adopt a cat!</h2>
<a href="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a href="/dogs">Visit dogs page</a>
```
