---
id: 68f1f993460f06a64c033061
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

L'attributo `poster` è un'immagine in miniatura del video. Pensa ai video che guardi su YouTube. Viene mostrata mentre il video si sta scaricando. Se l'attributo viene omesso, durante la fase di download viene mostrato il primo fotogramma del video.

Ora, aggiungi l'attributo `poster` con il valore `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg` al tuo elemento `video`.

# --hints--

Il tuo elemento `video` dovrebbe avere l'attributo `poster` con il valore `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('poster'), 'https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
--fcc-editable-region--
    
--fcc-editable-region--
  >
  </video>
</body>
</html>
```
