---
id: 68f1f8c248306da204655117
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

L'attributo `muted` silenzierà l'audio alla riproduzione iniziale. Se hai abilitato `controls`, l'utente potrà riattivare l'audio. Omettere l'attributo `muted` farà partire l'audio alla riproduzione iniziale.

L'attributo `muted` è un attributo booleano e non necessita di un valore.

Aggiungi l'attributo `muted` all'elemento `video`.

# --hints--

Il tuo elemento `video` dovrebbe avere l'attributo `muted`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('muted'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop controls>
--fcc-editable-region--
  </video>
</body>
</html>
```
