---
id: 68f1f6f97af9749a42260a43
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

L'attributo `controls` fornisce i controlli di riproduzione, inclusi riproduzione, riavvolgimento e controllo del volume per l'elemento `video`.

L'attributo `controls` è un attributo booleano e non necessita di un valore.

Aggiungi l'attributo `controls` all'elemento `video`.

Ora dovresti vedere l'elemento `video` visualizzato sulla pagina.

# --hints--

Il tuo elemento `video` dovrebbe avere l'attributo `controls`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop>
--fcc-editable-region--
  </video>
</body>
</html>
```
