---
id: 68f1f57723debf9383c75947
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

L'attributo `loop` riavvierà il video una volta terminata la riproduzione. Pensa a un meme di internet che ripete la riproduzione. Omettere l'attributo `loop` farà riprodurre il video una sola volta.

L'attributo `loop` è un attributo booleano e non necessita di un valore.

Aggiungi l'attributo `loop` all'elemento `video`.

# --hints--

Il tuo elemento `video` dovrebbe avere l'attributo `loop`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640">
--fcc-editable-region--
  </video>
</body>
</html>
```
