---
id: 68f1f371fefd448b4f70b898
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

In una lezione precedente, hai imparato i diversi attributi disponibili per l'elemento `video`. L'attributo `width` determina la larghezza del video in pixel.

Aggiungi l'attributo `width` all'elemento `video` con un valore di `640`.

# --hints--

Il tuo elemento `video` dovrebbe avere l'attributo `width` con il valore `640`.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('width'), '640');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video>
--fcc-editable-region--
  </video>
</body>
</html>
```
