---
id: 68f18724d71db1dda8d6fec7
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo workshop costruirai un lettore video HTML. Il codice boilerplate HTML ti è stato fornito.

Crea un elemento `h1` e assegnagli il testo `Working with the HTML Video Element`.

# --hints--

Dovresti avere un elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Il tuo elemento `h1` dovrebbe contenere il testo `Working with the HTML Video Element`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Working with the HTML Video Element');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
