---
id: 68f1fac259a4afabb478b960
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Avrai notato che non hai collegato il video vero e proprio. Lo farai nella fase successiva. Per quanto riguarda i tipi di file video, ci sono differenze nel supporto dei browser. Per gestire questo, puoi usare elementi `source` all'interno dell'elemento `video` e il browser selezionerà il primo `source` compatibile.

Ecco un esempio di elemento `source`:

```html
<video controls width="250">
  <source src="src-url-goes-here" type="video-type-goes-here" />
</video>
```

L'elemento `source` è un elemento vuoto quindi non ha un tag di chiusura.

Aggiungi un elemento `source` all'interno del tuo elemento `video`.


# --hints--

Dovresti avere un elemento `source` all'interno del tuo elemento `video`.

```js
assert.exists(document.querySelector('video > source'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
