---
id: 68f1eb9ad0a7575e4dbc31d5
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Successivamente, crea un elemento `video` sotto il `h1`. Nei prossimi passaggi aggiungerai gli attributi necessari per rendere funzionante il lettore video.

# --hints--

Dovresti avere un elemento `video`.

```js
assert.exists(document.querySelector('video'));
```

Il nuovo elemento `video` dovrebbe trovarsi sotto il `h1`.

```js
assert.exists(document.querySelector('h1 + video'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
