---
id: 68f26a5f79e97c5032c10642
title: Passaggio 10
challengeType: 0
dashedName: step-10
---

# --description--

Hai usato un file video con un'estensione `mp4` e devi indicare al browser quale sia, così sa come leggere il file.

Userai l'attributo `type` per specificare il tipo MIME `video/mp4`.

MIME (Multipurpose Internet Mail Extensions) è uno standard per descrivere documenti in formati diversi dal testo ASCII, per esempio audio, video e immagini.

MP4, formalmente noto come MPEG-4 Parte 14, è un formato contenitore multimediale digitale. È molto usato per memorizzare video e audio, ma può includere anche altri tipi di dati come sottotitoli e immagini statiche. I file MP4 sono progettati per lo streaming su Internet e sono compatibili con molti dispositivi e piattaforme.

Ora, aggiungi l'attributo `type` e il valore `video/mp4`.

# --hints--

Il tuo elemento `source` dovrebbe avere un attributo `type`.

```js
const source = document.querySelector('source');
assert.isTrue(source?.hasAttribute('type'));
```

Il tuo elemento `source` dovrebbe avere un attributo `type` con il valore `video/mp4`.

```js
const source = document.querySelector('source');
assert.strictEqual(source?.getAttribute('type'), 'video/mp4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
--fcc-editable-region--
      
--fcc-editable-region--
    >
  </video>
</body>
</html>
```
