---
id: 68f2701b2160c765f23b60f8
title: Passaggio 12
challengeType: 0
dashedName: step-12
---

# --description--

Un altro tipo MIME comune è il tipo MIME `video/ogg`.

Ogg è un formato contenitore multimediale digitale progettato per consentire lo streaming efficiente e la manipolazione di contenuti multimediali digitali. È mantenuto dalla Xiph.Org Foundation ed è libero e aperto, senza restrizioni di brevetti software. Il suo nome deriva da "ogging", gergo del gioco per computer Netrek.

Sotto il tuo secondo elemento `source`, aggiungi un terzo elemento `source` e dagli un attributo `src` con valore `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg` e un attributo `type` con valore `video/ogg`.

# --hints--

Dovresti avere un terzo elemento `source`.

```js
assert.exists(document.querySelector('video > source:nth-of-type(3)'));
```

Il tuo terzo elemento `source` dovrebbe avere un attributo `src`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('src'));
```

Il tuo terzo elemento `source` dovrebbe avere un attributo `src` con valore `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg');
```

Il tuo terzo elemento `source` dovrebbe avere un attributo `type`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('type'));
```

Il tuo terzo elemento `source` dovrebbe avere un attributo `type` con valore `video/ogg`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('type'), 'video/ogg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
