---
id: 68f26d52b13e0f5b6e2d2e09
title: Passaggio 11
challengeType: 0
dashedName: step-11
---

# --description--

Un altro tipo MIME comune è il tipo MIME `video/webm`.

WebM è un formato di file multimediale audiovisivo open source sviluppato da Google, progettato principalmente per contenuti multimediali basati sul web. Supporta codec video come VP8, VP9 e AV1, e codec audio come Vorbis e Opus, rendendolo una scelta popolare per gli elementi video e audio HTML5.

Sotto il tuo primo elemento `source`, aggiungi un altro elemento `source` e dagli un attributo `src` con valore `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm` e un attributo `type` con valore `video/webm`.

# --hints--

Dovresti avere un secondo elemento `source`.

```js
assert.exists(document.querySelector('video > source:nth-of-type(2)'));
```

Il tuo secondo elemento `source` dovrebbe avere un attributo `src`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('src'));
```

Il tuo secondo elemento `source` dovrebbe avere un attributo `src` con valore `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm');
```

Il tuo secondo elemento `source` dovrebbe avere un attributo `type`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('type'));
```

Il tuo secondo elemento `source` dovrebbe avere un attributo `type` con valore `video/webm`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('type'), 'video/webm');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
