---
id: 68f27139bddf9f6bf2821acf
title: Passaggio 13
challengeType: 0
dashedName: step-13
---

# --description--

L'ultimo elemento `source` che aggiungerai sarà per il tipo MIME `video/quicktime`.

QuickTime è un'architettura multimediale estensibile creata da Apple, che supporta la riproduzione, lo streaming, la codifica e la transcodifica di vari formati di media digitali. Non è popolare come il formato MP4, ma potresti averne bisogno per il supporto di applicazioni legacy.

Sotto il tuo terzo elemento `source`, aggiungi un quarto elemento `source` e dagli un attributo `src` con valore `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov` e un attributo `type` con valore `video/quicktime`.

Congratulazioni! Hai completato il workshop sul lettore video HTML.

# --hints--

Dovresti avere un quarto elemento `source`.

```js
assert.exists(document.querySelector('video > source:nth-of-type(4)'));
```

Il tuo quarto elemento `source` dovrebbe avere un attributo `src`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('src'));
```

Il tuo quarto elemento `source` dovrebbe avere un attributo `src` con valore `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov');
```

Il tuo quarto elemento `source` dovrebbe avere un attributo `type`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('type'));
```

Il tuo quarto elemento `source` dovrebbe avere un attributo `type` con valore `video/quicktime`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('type'), 'video/quicktime');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    controls
    loop
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
       src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov"
      type="video/quicktime"
    >
  </video>
</body>
</html>
```
