---
id: 6853e58d52ec64349a5760ca
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este taller, implementarás el algoritmo de camino más corto. Escribirás una función en Python que calcule el camino más corto entre los nodos en un grafo, y también devuelva el camino recorrido.

Por ejemplo, dado un grafo donde las ciudades están conectadas por carreteras con diferentes distancias, el algoritmo encontrará la ruta más corta de una ciudad a otra. Si quieres viajar de City A a City D, el algoritmo podría encontrar que ir A ⇨ B ⇨ C ⇨ D (total: 15km) es más corto que ir directamente A ⇨ D (20km).

Para comenzar, define una variable llamada `INF` y asígnale el valor `float('inf')`, que representa infinito positivo. Más adelante, la usarás para indicar una distancia infinita entre dos nodos.

# --hints--

Debes definir la variable `INF` en la parte superior de tu archivo.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("INF")`));
  }
});
```

Debes asignar `float('inf')` a la variable `INF`.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("INF").is_equivalent("INF = float('inf')")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
