---
id: 686253f1b051998ad4904e3e
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Ahora crearás la función principal que acepta tres parámetros: la matriz de adyacencia, el nodo inicial y un nodo objetivo opcional.

Crea una función llamada `shortest_path` que reciba tres parámetros: `matrix`, `start_node` y `target_node`. Asigna `None` como valor predeterminado para `target_node`.

El parámetro `target_node` está establecido en `None` por defecto, lo que indica que si no se especifica un nodo objetivo, la función debe calcular las rutas más cortas desde el nodo inicial a todos los demás nodos en el grafo.

Agrega una sentencia `pass` dentro del cuerpo de la función por ahora.

# --hints--

Debes crear una función llamada `shortest_path`.

```js
({
    test: () => {
      assert(runPython(`_Node(_code).has_function("shortest_path")`));
    }
});
```

La función debe recibir tres parámetros: `matrix`, `start_node` y `target_node`, siendo `target_node` un valor predeterminado de `None`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_args("matrix, start_node, target_node=None")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--

--fcc-editable-region--
```
