---
id: 686251fb49ebc489ba60ccb5
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Necesitarás crear una lista 2D para representar la matriz de adyacencia del grafo. Esta matriz se usará para representar los pesos de las aristas entre los nodos del grafo.

Crea una variable llamada `adj_matrix` y asígnale una lista 2D que represente el grafo con los siguientes pesos:

```py
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

Debes crear una variable llamada `adj_matrix`.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("adj_matrix")`));
  }
});
```

`adj_matrix` debe ser una lista 2D que contenga los valores proporcionados.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("adj_matrix").is_equivalent("adj_matrix = [[0, 5, 3, INF, 11, INF], [5, 0, 1, INF, INF, 2], [3, 1, 0, 1, 5, INF], [INF, INF, 1, 0, 9, 3], [11, INF, 5, 9, 0, INF], [INF, 2, INF, 3, INF, 0]]")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
--fcc-editable-region--

--fcc-editable-region--
```
