---
id: 68625ad34aa3f58d7f3b9b07
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Necesitas almacenar el número de nodos en el grafo. Para esto necesitarás una variable que coincida con la longitud de la matriz de adyacencia.

Dentro de la función `shortest_path`, crea una variable `n` y asígnale la longitud de la `matrix`.

# --hints--

Debes inicializar `n` con la longitud de la `matrix`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("n").is_equivalent("n = len(matrix)")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    pass

--fcc-editable-region--
```
