---
id: 686b73007fc582db0132c361
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Necesitas llevar un registro de la distancia conocida más corta desde el nodo inicial hasta cada otro nodo en el grafo.

Para hacer esto, crea una variable llamada `distances` e inicialízala como una lista que contiene un solo elemento: `INF`.  


# --hints--

Debes crear una variable llamada `distances`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("distances")`
        )
      );
    }
});
```

La lista `distances` debe inicializarse con `[INF]`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF]")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)

--fcc-editable-region--
```
