---
id: 686b78bc13d518e1267448e4
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

Además de rastrear distancias, también necesitas llevar un registro de las rutas reales tomadas para llegar a cada nodo.

Crearás una lista donde cada entrada almacena el camino tomado para llegar a ese nodo. Inicialmente, el camino de cada nodo solo contendrá a sí mismo.  

Las comprensiones de listas ofrecen una forma concisa de crear listas. Por ejemplo:

```py
[x * 2 for x in range(3)]
```

Crea una variable llamada `paths` e inicialízala usando una comprensión de listas que crea una lista que contiene `[node_no]` para cada `node_no` en `range(n)`.

# --hints--

Debes tener una variable llamada `paths`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("paths")`
        )
      );
    }
});
```

Debes inicializar `paths` usando una comprensión de listas.

```js
({  
    test: () => {  
      runPython(`  
import ast  
assert isinstance(_Node(_code).find_function("shortest_path").find_variable("paths").tree.value, ast.ListComp)  
    `);  
    }  
});
```

Tu comprensión de listas debe usar `node_no` para iterar sobre `range(n)`.

```js
({  
  test: () => {  
    runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert len(v.find_comp_targets()) == 1  
assert v.find_comp_targets()[0].is_equivalent("node_no")  
assert len(v.find_comp_iters()) == 1  
assert v.find_comp_iters()[0].is_equivalent("range(n)")  
    `);  
  }  
});
```

Tu comprensión de listas debe evaluar `[node_no]` para cada `node_no` en `range(n)`.  

```js  
({  
    test: () => {  
      runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert v.find_comp_expr().is_equivalent("[node_no]")  
    `);  
    }  
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0

--fcc-editable-region--
```
