---
id: 686b82af323098e7a5203bf4
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

Mientras el algoritmo se ejecuta, necesitas llevar un registro de qué nodos ya has visitado, para no procesarlos más de una vez.

Para hacer esto, crea una lista llamada `visited` e inicialízala con `False` para cada nodo.

# --hints--

Debes tener una lista llamada `visited` dentro de la función `shortest_path`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("visited")`
        )
      );
    }
});
```

Debes inicializar la lista `visited` con `False` para cada nodo usando `[False] * n`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("visited").is_equivalent("visited = [False] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]

--fcc-editable-region--
```
