---
id: 6876089f7a500b7c782e5b83
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

Necesitas decidir si el nodo actual es una mejor opción que el que ya has encontrado (si es que hay alguno). Para hacer esto, agregarás una sentencia condicional dentro del ciclo.

La condición debe hacer dos cosas:

- Verifica si el nodo no ha sido visitado aún.
- Compara la distancia conocida con la `min_distance` actual.

Dentro del `for` interno, agrega una sentencia `if` que verifique si `node_no` no ha sido visitado **y** si `distances[node_no]` es menor que `min_distance`.

Agrega `pass` como un marcador de posición dentro del bloque condicional.

# --hints--

Debes agregar una estructura `if` dentro de tu `for` interno.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0]`
        )
      );
    }
});
```

Tu sentencia `if` debe verificar si el nodo no ha sido visitado (`visited[node_no]` es falso) y `distances[node_no]` es menor que `min_distance`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("not visited[node_no] and distances[node_no] < min_distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

--fcc-editable-region--
    for _ in range(n):  
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            pass

--fcc-editable-region--
```
