---
id: 68760671d950767b055bf2f9
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Ahora necesitas revisar cada nodo para encontrar el que tiene la distancia conocida más pequeña que aún no ha sido visitado.

Para hacer esto, agrega un ciclo `for` dentro del ciclo principal.

El ciclo debe iterar a través de cada `node_no` en `range(n)`, donde `n` es el número de nodos en el grafo. Añade `pass` como el cuerpo del ciclo por ahora.

# --hints--

Debes tener un segundo `for` dentro de la función `shortest_path`.

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
len(outer_loop.find_bodies()[0].find_for_loops()) >= 1
      `)
      );
    }
});
```

Tu ciclo `for` debe iterar sobre `range(n)` con `node_no` como la variable de iteración.

```js
({
    test: () => {
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
inner_loop = outer_loop.find_bodies()[0].find_for_loops()[0]
assert inner_loop.find_for_vars().is_equivalent("node_no")
assert inner_loop.find_for_iter().is_equivalent("range(n)")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

--fcc-editable-region--
    for _ in range(n):  
        min_distance = INF  
        current = -1

--fcc-editable-region--
```
