---
id: 68760464eb9c2e79cc913fac
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

En este paso, agregarás un ciclo que se ejecutará una vez por cada nodo en el grafo. Este ciclo permitirá que el algoritmo actualice distancias y caminos a lo largo de múltiples iteraciones.

Crea un ciclo `for` que se ejecute `n` veces. Usa `_` como la variable del ciclo ya que no necesitas usar el valor de la iteración.  

Dentro del ciclo, necesitas prepararte para seleccionar el siguiente nodo a procesar creando dos variables: 

- uno para almacenar la distancia más pequeña encontrada hasta ahora en la iteración actual
- y otro para almacenar el índice del nodo que tiene esta distancia más pequeña.

Crea las variables `min_distance` y `current`, y asígnalas a `INF` y `-1`, respectivamente. 

# --hints--

Debes tener un ciclo `for` dentro de la función `shortest_path`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0]`
        )
      );
    }
});
```

Tu ciclo `for` debe iterar sobre `range(n)`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_iter().is_equivalent("range(n)")`
        )
      );
    }
});
```

Tu ciclo for debe usar `_` como la variable de iteración.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_vars().is_equivalent("_")`
        )
      );
    }
});
```

Debes crear una variable llamada `min_distance` y asignarle `INF` dentro del ciclo for.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("min_distance = INF")`
        )
      );
    }
});
```

Debes crear una variable llamada `current` y asignarle `-1` dentro del ciclo for.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("current = -1")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
--fcc-editable-region--

--fcc-editable-region--
```
