---
id: 687774e17e0972345de2527a
title: Paso 15
challengeType: 20
dashedName: step-15
---

# --description--

Si se encontró un nodo válido en la pasada, necesitas marcarlo como visitado para que no se considere nuevamente en iteraciones futuras.

Después de la sentencia `if` que agregaste en el paso anterior, asigna `True` a `visited[current]`.

# --hints--

Debes establecer `visited[current]` en `True` después de tu última sentencia `if`.

```js
({
    test: () => {
      runPython(`
      block = _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0]
      if_stmt = """if current == -1:
        break"""
      target_stmt = "visited[current] = True"
      assert block.is_ordered(if_stmt, target_stmt)
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            if not visited[node_no] and distances[node_no] < min_distance:  
                min_distance = distances[node_no]  
                current = node_no

--fcc-editable-region--
        if current == -1:  
            break

--fcc-editable-region--
```
