---
id: 68777816e95d4936f129819f
title: Paso 17
challengeType: 20
dashedName: step-17
---

# --description--

Antes de intentar actualizar la distancia a un vecino, necesitas verificar que el vecino sea accesible y no visitado. Luego calcularás cuál sería la distancia total para llegar a ese vecino a través del nodo actual.

Dentro del ciclo `for`, agrega una sentencia `if` que verifique:

- La `distance` no es igual a `INF` (lo que significa que hay una arista entre los nodos)
- El vecino `node_no` aún no ha sido visitado

Dentro del bloque `if`, crea una variable llamada `new_distance` y asígnale la suma de `distances[current]` (la distancia más corta al nodo actual) y `distance` (la distancia desde el nodo actual hasta el vecino).

# --hints--

Debes tener una estructura `if` que verifique si `distance` no es igual a `INF` y si el nodo vecino no ha sido visitado.

```js
({
    test: () => {
      assert(
        runPython(`_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("distance != INF and not visited[node_no]") 
      `)
      );
    }
});
```

Debes crear una variable llamada `new_distance` dentro del bloque `if`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("new_distance")`
        )
      );
    }
});
```

Debes asignar a `new_distance` la suma de `distances[current]` y `distance`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("new_distance").is_equivalent("new_distance = distances[current] + distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
              min_distance = distances[node_no]
              current = node_no

        if current == -1:
            break
        visited[current] = True

--fcc-editable-region--
        for node_no in range(n):  
            distance = matrix[current][node_no]

--fcc-editable-region--
```
