---
id: 687b3926419dec576850b814
title: Paso 18
challengeType: 20
dashedName: step-18
---

# --description--

Ahora que has calculado la nueva distancia posible al vecino, verifica si es mejor que la que está almacenada actualmente en la lista `distances`. Si lo es, actualiza la distancia.

Dentro del bloque `if` existente, agrega una sentencia `if` que verifique si `new_distance` es menor que `distances[node_no]`.

Dentro de este nuevo bloque condicional, actualiza `distances[node_no]` para almacenar el `new_distance`.

# --hints--

Deberías tener una estructura `if` anidada que verifique si `new_distance` es menor que `distances[node_no]`.

```js
({
    test: () => {
      runPython(`  
      cond = _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]  
      assert cond.is_equivalent("new_distance < distances[node_no]") or cond.is_equivalent("distances[node_no] > new_distance")  
      `);
    }
});
```

Debes actualizar `distances[node_no]` para almacenar la `new_distance`.

```js
({
    test: () => {
      assert(
        runPython(`_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("distances[node_no] = new_distance")
      `)
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            if not visited[node_no] and distances[node_no] < min_distance:  
                min_distance = distances[node_no]  
                current = node_no

        if current == -1:  
            break  
        visited[current] = True  

--fcc-editable-region--  
        for node_no in range(n):
            distance = matrix[current][node_no]  
            if distance != INF and not visited[node_no]:  
                new_distance = distances[current] + distance

--fcc-editable-region--
```
