---
id: 68e4c0ff9edcc92771d7549a
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

En este momento, la lista `distances` solo tiene un elemento. Sin embargo, necesitas un valor de distancia para cada nodo en el grafo.

En Python, puedes multiplicar una lista por un entero para repetirla. Por ejemplo, `[0] * 3` crea `[0, 0, 0]`.

Actualiza la lista `distances` para que contenga `n` copias de `INF` multiplicando `[INF]` por `n`.

# --hints--

Debes actualizar `distances` para que contenga `n` copias de `INF`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF]

--fcc-editable-region--
```
