---
id: 69757cc0faae0152c1418aad
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo workshop, eserciterai l’uso dei numeri e delle operazioni matematiche per creare un divisore del conto. Questo strumento calcolerà quanto deve pagare ciascuna persona dopo aver aggiunto i costi del pasto e la mancia.

Per cominciare, ti serve un modo per tenere traccia dell’importo totale mentre aggiungi i costi. In Python, puoi usare una variabile per archiviare un intero (un numero intero) che cambia nel tempo.

Per esempio, potresti scrivere:

```py
my_number = 2
```

Crea una variabile chiamata `running_total` e assegnale il valore `0`.

# --hints--

Dovresti avere una variabile chiamata `running_total`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('running_total')
    `))
})
```

Dovresti assegnare l’intero `0` alla tua variabile `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('running_total').is_equivalent('running_total = 0')
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
