---
id: 6976d8a7cd2a1a5ffcd7f9c4
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Ogni corso ha un costo. Devi archiviare questi importi in variabili per poterli usare in seguito. Poiché questi importi includono i centesimi, userai il tipo float, che serve a rappresentare numeri decimali. Ecco un esempio di variabile con un valore float:

```py
change = 2.35
```

Crea quattro variabili: `appetizers` impostata su `37.89`, `main_courses` impostata su `57.34`, `desserts` impostata su `39.39` e `drinks` impostata su `64.21`. 

# --hints--

Dovresti definire una variabile chiamata `appetizers` e assegnarle il valore `37.89`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('appetizers').is_equivalent('appetizers = 37.89')
    `))
})
```

Dovresti definire una variabile chiamata `main_courses` e assegnarle il valore `57.34`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('main_courses').is_equivalent('main_courses = 57.34')
    `))
})
```

Dovresti definire una variabile chiamata `desserts` e assegnarle il valore `39.39`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('desserts').is_equivalent('desserts = 39.39')
    `))
})
```

Dovresti definire una variabile chiamata `drinks` e assegnarle il valore `64.21`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('drinks').is_equivalent('drinks = 64.21')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

--fcc-editable-region--

--fcc-editable-region--
```
