---
id: 6976cc0b0c135686a4620b94
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Ora devi considerare il numero di persone che dividono il conto. Archivia questo valore in una variabile, come hai fatto nel passaggio precedente.

Crea una variabile chiamata `num_of_friends` e assegnale il valore di `4`. Questa verrà usata più avanti nel workshop per calcolare la divisione finale.

# --hints--

Dovresti avere una variabile chiamata `num_of_friends`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('num_of_friends')
    `))
})
```

Dovresti assegnare l'intero `4` alla tua variabile `num_of_friends`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('num_of_friends').is_equivalent('num_of_friends = 4')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

--fcc-editable-region--

--fcc-editable-region--
```
