---
id: 6977a4306bb3e856690a4890
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Ora che hai calcolato la mancia, devi aggiungerla al tuo `running_total` per trovare l'importo finale del conto.

Usa l'operatore `+=` per aggiungere il valore di `tip` al tuo `running_total`. Infine, usa `print()` per stampare la stringa `Total with tip:` seguita da uno spazio e dal valore di `running_total`.

# --hints--

Dovresti usare l'operatore `+=` per aggiungere il valore di `tip` al tuo `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_stmt('running_total += tip')
    `))
})
```

Dovresti stampare la stringa `Total with tip:` seguita da uno spazio e dalla variabile `running_total`.

```js
({
    test: () => runPython(`
import io, contextlib, re

buffer = io.StringIO()
with contextlib.redirect_stdout(buffer):
    exec(_code)

match = re.search(r"Total with tip: ([0-9]+(?:\\.[0-9]+)?)", buffer.getvalue())

assert match
assert abs(float(match.group(1)) - ((37.89 + 57.34 + 39.39 + 64.21) * 1.25)) < 1e-6
    `)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

--fcc-editable-region--

--fcc-editable-region--
```
