---
id: 6977a4306bb3e856690a4890
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Ora che hai calcolato la mancia, devi aggiungerla al tuo `running_total` per trovare l'importo finale del conto.

In Python, puoi usare l'operatore di assegnazione aumentata `+=` per aggiungere un valore a una variabile e aggiornare quella variabile allo stesso tempo. Per esempio, `total += 5` è un modo abbreviato per scrivere `total = total + 5`.

Usa l'operatore `+=` per aggiungere il valore di `tip` al tuo `running_total`. Infine, usa `print()` per stampare la stringa `Total with tip:` seguita da uno spazio e dal valore di `running_total`.

# --hints--

Dovresti usare l'operatore `+=` per aggiungere il valore di `tip` al tuo `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_stmt('running_total += tip')
    `))
})
```

Dovresti stampare la stringa `Total with tip:` seguita da uno spazio e dalla variabile `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Total with tip:', running_total)") or _Node(_code).has_call("print(f'Total with tip: {running_total}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

--fcc-editable-region--

--fcc-editable-region--
```
