---
id: 6977a8bbcbdd557886bd0718
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

Ora che la mancia è inclusa, hai l'importo finale per tutto il gruppo. Devi determinare quanto deve pagare ciascuna persona dividendo il conto totale per il numero di amici.

In Python, usi la barra `/` per eseguire la divisione. Per esempio:

```py
half = 10 / 2
```

Crea una variabile chiamata `final_bill` e assegnale il risultato della divisione di `running_total` per `num_of_friends`.

Infine, usa la funzione `print()` per stampare la stringa `Bill per person:` seguita da uno spazio e dal valore di `final_bill`.

# --hints--

Dovresti definire una variabile chiamata `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('final_bill')
    `))
})
```

Dovresti usare l'operatore `/` per dividere `running_total` per `num_of_friends` e assegnare il risultato alla tua variabile `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('final_bill').is_equivalent('final_bill = running_total / num_of_friends')
    `))
})
```

Dovresti stampare la stringa `Bill per person:` seguita da uno spazio e dalla variabile `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Bill per person:', final_bill)") or _Node(_code).has_call("print(f'Bill per person: {final_bill}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

running_total += tip
print('Total with tip:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
