---
id: 69779dcb44016eccc05c15a4
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Il servizio è stato eccellente, quindi il gruppo decide di lasciare una mancia del 25%. Per calcolare una percentuale in Python, puoi moltiplicare il totale per l'equivalente decimale della percentuale.

Per esempio, per trovare il 10% di un valore, lo moltiplicheresti per `0.10` usando l'operatore `*`:

```py
tax = total * 0.10
```

Crea una variabile chiamata `tip` e assegnale il risultato della moltiplicazione di `running_total` per `0.25`.

Infine, usa `print()` per stampare la stringa `Tip amount:` seguita da uno spazio e dal valore della tua variabile `tip`.

# --hints--

Dovresti avere una variabile chiamata `tip`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('tip')
    `))
})
```

La tua variabile `tip` dovrebbe essere il risultato di `running_total * 0.25`.

```js
({
    test: () => runPython(`
    t = _Node(_code).find_variable('tip')
    assert t.is_equivalent('tip = running_total * 0.25') or t.is_equivalent('tip = 0.25 * running_total')`)
})
```

Dovresti stampare la stringa `Tip amount:` seguita da uno spazio e dalla variabile `tip`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Tip amount:', tip)") or _Node(_code).has_call("print(f'Tip amount: {tip}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
