---
id: 6982684f3a25f379e195a5fc
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Ora che hai archiviato i costi individuali, puoi calcolare il totale.

Ricorda che l'operatore `+=` aggiunge un valore a una variabile esistente e la aggiorna contemporaneamente. Per esempio:

```py
total = 10
total += 2 + 2 + 1
print(total)  # total is now 15
```

Usa l'operatore `+=` una volta per aggiungere `appetizers`, `main_courses`, `desserts` e `drinks` a `running_total`.

Infine, usa `print()` per mostrare la stringa `Total bill so far:` seguita da uno spazio e dal valore di `running_total`.

**Nota:** Potresti notare che l'output ha più cifre decimali del previsto. Come hai imparato in una lezione precedente, questo succede perché i numeri sono archiviati in binario e molti valori decimali non possono essere rappresentati esattamente in questo formato, il che porta a errori di arrotondamento.

# --hints--

Dovresti aggiornare `running_total` usando l'operatore `+=` una volta per aggiungere tutti e quattro i costi del corso.

```js
({
       test: () => runPython(`
import itertools  
perms = itertools.permutations(['+ appetizers', '+ main_courses', '+ desserts', '+ drinks'])  
values = (' '.join(perm).lstrip('+') for perm in perms)  
solutions = (f'running_total += {v}' for v in values)  
assert any(_Node(_code).has_stmt(s) for s in solutions)`)
})
```

Dovresti stampare la stringa `Total bill so far:` seguita da uno spazio e dal valore di `running_total`.

```js
({
    test: () => runPython(`
import io, contextlib

_expected = f"Total bill so far: {37.89 + 57.34 + 39.39 + 64.21}"

buffer = io.StringIO()
with contextlib.redirect_stdout(buffer):
    exec(_code)

assert _expected in buffer.getvalue()
`)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

--fcc-editable-region--

--fcc-editable-region--
```
