---
id: 660ef6355e8f5a9e67fe5f46
title: Paso 16
challengeType: 1
dashedName: step-16
---

# --description--

Con el tipo de dato `number`, puedes realizar operaciones matemáticas, como la adición. Intenta imprimir `count + 1` en la consola.

# --hints--

Debes acceder a la `consola` en tu código.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

Debes acceder al `.log()` de la `consola`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Debes imprimir `count + 1` en la consola.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*count\s*\+\s*1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
