---
id: 660ef55dd468079679ee0092
title: Paso 12
challengeType: 1
dashedName: step-12
---

# --description--

Ahora estás listo para declarar tu siguiente variable. Elimina ambos `console.log` y la reasignación `character`.

También elimina tu variable `secondCharacter`, pero deja la inicialización de la variable `character` intacta.

# --hints--

No deberías tener ningún `console.log()` en tu código.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

No deberías reasignar el valor de `character`.

```js
assert.isAtMost(__helpers.removeJSComments(code).match(/character\s*=/g).length, 1);
```

Tu variable `character` debería tener la cadena de texto `"Hello"` como valor.

```js
assert.equal(character, "Hello");
```

No deberías tener una variable `secondCharacter`.

```js
assert.notMatch(__helpers.removeJSComments(code), /secondCharacter/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
secondCharacter = character;
console.log(secondCharacter);
--fcc-editable-region--
```
