---
id: 660ef31a5be625914a0102cd
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

Cuando una variable es declarada con la palabra reservada `let`, puedes <dfn>reasignar</dfn> (o cambiar) el valor de esa variable más adelante. En este ejemplo, el valor de `programmer` se cambia de `"Naomi"` a `"CamperChan"`.

```js
let programmer = "Naomi";
programmer = "CamperChan";
```

Ten en cuenta que al reasignar una variable, **no** se utiliza la palabra clave `let` nuevamente.

Después de tu `console.log`, asigna el valor `"World"` a tu variable `character`.

# --hints--

Debes usar `character` en tu código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character/g), 3);
```

Debes usar el operador de asignación para reasignar `character`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\s*\=/g), 2);
```

Debes usar la cadena de texto `"World"` en tu código.

```js
assert.match(__helpers.removeJSComments(code), /("|')World\1/);
```

Tu variable `character` debería tener el valor `"World"`.

```js
assert.equal(character, "World");
```

Tu reasignación no debe usar `let`.

```js
assert.isBelow(__helpers.removeJSComments(code).match(/let/g).length, 2);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);

--fcc-editable-region--
```
