---
id: 660ef857f2806aa626d29d17
title: Paso 17
challengeType: 1
dashedName: step-17
---

# --description--

También puedes realizar restas (`-`), multiplicación (`*`), y división (`/`). Siéntete libre de experimentar con los operadores y números en tu `console.log`. Cuando estés listo para avanzar, elimina el `console.log`.

# --hints--

No deberías tener ningún `console.log()` en tu código.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Tu variable `character` debe seguir teniendo el valor `"Hello"`.

```js
assert.equal(character, "Hello");
```

Tu variable `count` debe seguir teniendo el valor de `8`.

```js
assert.equal(count, 8);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let count = 8;
console.log(count + 1);
--fcc-editable-region--
```
