---
id: 660f1b6e60bd9edf902c81fd
title: Paso 36
challengeType: 1
dashedName: step-36
---

# --description--

Tu instrucción <dfn>iteration</dfn> le dirá a tu bucle qué hacer con el iterador después de cada ejecución.

Cuando reasigna una variable, puede utilizar la variable para referenciar el valor anterior antes de la reasignación. Esto te permite hacer cosas como añadir tres a un número existente. Por ejemplo, `bees = bees + 3` incrementará el valor de `bees` en tres.

Usa esa sintaxis para reemplazar tu cadena de texto `"iteration"` con una sentencia de reasignación que aumenta `i` por uno.

# --hints--

Deberías añadirle 1 a tu variable `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\s*1/);
```

Debes asignar `i + 1` de vuelta a tu variable `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

Tu bucle `for` debe aumentar `i` en `1` en cada iteración.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; "iteration") {

}
--fcc-editable-region--
```
