---
id: 660f2fbd45b520046cac68e8
title: Paso 65
challengeType: 1
dashedName: step-65
---

# --description--

Recuerda que en un paso anterior, aprendiste sobre los valores devueltos. Una función puede <dfn>devolver</dfn> un valor para que tu aplicación lo consuma por separado.

En una función, la palabra reservada `return` se utiliza para especificar un valor de retorno. Por ejemplo, esta función devolvería el valor dado al primer parámetro:

```js
function name(parameter) {
  return parameter;
}
```

Utiliza la palabra reservada `return` para devolver el valor de la variable `character`, repetido `rowNumber` veces.

# --hints--

Debes utilizar el método `.repeat()`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/\.repeat\(/g), 2);
```

Deberías utilizar el método `.repeat()` en tu constante `character`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\.repeat\(/g), 2);
```

Deberías pasar `rowNumber` a tu llamada `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*rowNumber\s*\)/);
```

Debes usar la palabra reservada `return`.

```js
assert.match(__helpers.removeJSComments(code), /return/);
```

Deberías devolver el resultado de tu llamada `.repeat()`.

```js
assert.equal(padRow(3), "###");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
