---
id: 660f3915b41a441537ec9f5e
title: Paso 72
challengeType: 1
dashedName: step-72
---

# --description--

Puesto que solo aumentas `i` en `1`, puedes usar el <dfn>operador de incremento</dfn> `++`. Este operador incrementa el valor de una variable en 1, actualizando la asignación para esa variable. Por ejemplo, `test` se convertiría en `8` aquí:

```js
let test = 7;
test++;
```

Reemplaza tu operador de asignación con el operador de incremento para la iteración de tu bucle.

# --hints--

Tu bucle `for` no debería usar el operador de asignación con `i`.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

Tu bucle `for` debe usar el operador de incremento en `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i += 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
