---
id: 660f455b044d3230ed971e98
title: Paso 85
challengeType: 1
dashedName: step-85
---

# --description--

Un bucle <dfn>`while`</dfn> se ejecutará una y otra vez hasta que la `condición` especificada deje de ser cierta. Tiene la siguiente sintaxis:

```js
while (condition) {
  logic;
}
```

Usa esa sintaxis para declarar un bucle `while` con `continueLoop` como condición. El cuerpo debe estar vacío.

# --hints--

Debes usar un bucle `while`.

```js
assert.match(__helpers.removeJSComments(code), /while/);
```

Tu bucle `while` debe usar `continueLoop` como condición.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
