---
id: 660f1a00ac619ddc1e259a66
title: Paso 35
challengeType: 1
dashedName: step-35
---

# --description--

La <dfn>condición</dfn> de un bucle `for` le dice al bucle cuántas veces debe iterar. Cuando la `condition` es falsa, el bucle se detendrá.

En JavaScript, un valor booleano puede ser `true` o `false`. Estas no son cadenas de texto; aprenderás más sobre la diferencia más adelante.

Por ahora, utilizarás el operador <dfn>menor que</dfn> (`<`). Esto te permite verificar si el valor de la izquierda es menor que el valor de la derecha. Por ejemplo, `count < 3` se evaluaría como `true` si `count` es `2` y `false` si `count` es `4`.

Reemplaza tu cadena de texto `"condition"` con una condición para verificar si `i` es menor que `count` (sin dejar espacios).

# --hints--

Debes utilizar el operador menor que.

```js
assert.match(__helpers.removeJSComments(code), /</);
```

Debes utilizar el operador menor que para comprobar si `i` es menor que `count`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

Tu bucle `for` debe usar `i < count` como condición.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*count\s*;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; "condition"; "iteration") {

}
--fcc-editable-region--
```
