---
id: 660f4cffb1459d45e34902d1
title: Paso 106
challengeType: 1
dashedName: step-106
---

# --description--

En realidad puedes construir la piramida invertida sin necesidad de hacer un bucle "hacia atrás" como lo hiciste.

Para hacer esto, necesitarás aprender un par de nuevos métodos de arrays. Empieza usando `const` para declarar una variable `numbers`. Asigna a un array los elementos `1`, `2` y `3`. Luego muestra en consola el array `numbers`.

# --hints--

Debes usar `const` para declarar una variable `numbers`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+numbers/);
```

Tu variable `numbers` debe ser un array.

```js
assert.isArray(numbers);
```

Tu array `numbers` debe tener los elementos `1`, `2`y `3` en ese orden.

```js
assert.deepEqual(numbers, [1,2,3]);
```

Deberías mostrar en pantalla tu array `numbers`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*numbers\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
