---
id: 660f359af3e32e0f1a6880b7
title: Paso 68
challengeType: 1
dashedName: step-68
---

# --description--

Ahora deberías ver el mismo grupo de caracteres en tu consola. Tu función `padRow` está haciendo exactamente lo mismo que hacías antes, pero ahora está en una sección reutilizable propia.

Utiliza el operador de adición para concatenar un espacio en blanco `" "` al principio y al final de tu cadena de texto repetida `character`.

Recuerda que puedes usar el operador `+` para concatenar cadenas como esta:

```js
" " + "string"
```

# --hints--

Debes concatenar un espacio vacío al principio del valor devuelto.

```js
assert.match(padRow(1, 1), /^\s/);
```

Debes concatenar un espacio vacío al final del valor devuelto.

```js
assert.match(padRow(1, 1), /\s$/);
```

Tu función `padRow()` debería devolver la serie `character` repetida con un espacio antes y después de la serie.

```js
assert.equal(padRow(1, 1), " # ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
