---
id: 660f383d4c772c12ff59904b
title: Paso 70
challengeType: 1
dashedName: step-70
---

# --description--

Puedes pasar expresiones completas como argumentos. La función recibirá el resultado de la evaluación de esa expresión. For example, these two function calls below would yield the same result. (Note that the order of operations rule PEMDAS—Parenthesis, Exponents, Multiplication, Division, Addition, Subtraction—applies):

```js
test(2 * 3 + 1);
test(7);
```

Volviendo a ver el patrón:

```md
··#··
·###·
#####
```

Actualiza el valor `character` para que se repita `2 * rowNumber - 1` veces.

Abra la consola de nuevo para ver el resultado actualizado.

# --hints--

Tu función debería pasar `2 * rowNumber - 1` a tu llamada `.repeat()`.

```js
assert.equal(padRow(4, 5), " ####### ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(rowNumber) + " ".repeat(rowCount - rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
