---
id: 660f17d4e9f227d86e834abd
title: Paso 33
challengeType: 1
dashedName: step-33
---

# --description--

Para generar una pirámide, necesitarás crear múltiples filas. Cuando tienes que realizar una tarea repetidamente hasta que se cumpla una condición, utilizarás un <dfn>bucle</dfn>. Hay muchas maneras de escribir un bucle.

Estás por comenzar con un bucle básico `for`. El bucle `for` usa la siguiente sintaxis:

```js
for (iterator; condition; iteration) {
  logic;
}
```

En los próximos pasos, explorarás cada componente del bucle en detalle. Por ahora, construye un bucle `for` que incluya los términos `"iterator"`, `"condition"` y `"iteration"` para los tres componentes. Mantén el <dfn>cuerpo</dfn> del bucle, la sección entre llaves `{}`, vacía.


# --hints--

Deberías tener un bucle `for`.

```js
assert.match(__helpers.removeJSComments(code), /for/);
```

El primer componente de tu bucle `for` debe ser la cadena de texto `"iterator"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1/);
```

El segundo componente de tu bucle `for` debe ser la cadena de texto `"condition"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2/);
```

El tercer componente de tu bucle `for` debe ser la cadena de texto `"iteration"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)/);
```

El cuerpo de tu bucle `for` debe estar vacío.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--
```
