---
id: 660f255022991ef34ed0ee88
title: Paso 44
challengeType: 1
dashedName: step-44
---

# --description--

La impresión de números no resultará en una pirámides visualmente atractiva. Ahora que está saliendo el contenido formateado de tu array `rows`, es hora de actualizar tu bucle original.

En lugar de "pushear" `i` al array, "pushea" el valor de tu constante `character`.

# --hints--

Ya no deberías "pushear" tu variable `i`.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

Deberías "pushear" tu variable `character`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
