---
id: 660f229d2dbe09ef2954a4a1
title: Paso 42
challengeType: 1
dashedName: step-42
---

# --description--

Recuerda que en tu bucle anterior utilizaste el operador de adición `+` para aumentar el valor de `i` por `1`.

Puedes hacer algo similar con un valor de cadena de texto, añadiendo una nueva cadena de texto a una cadena de texto existente. Por ejemplo, `hello = hello + " World";` añadiría la cadena de texto `" World"` a la cadena de texto existente almacenada en la variable `hello`. Esto se llama <dfn>concatenación</dfn>.

En tu bucle `for...of`, usa el operador de adición para concatenar el valor de `row` al valor de tu variable `result`.

# --hints--

Debes usar el operador de concatenación en tu variable `result`.

```js
assert.match(__helpers.removeJSComments(code), /(?:result\s*\+|\+\s*result)/);
```

Deberías concatenar `row` a tu variable `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row|row\s*\+\s*result/);
```

Debes asignar el resultado de tu concatenación de vuelta a la variable `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*(result\s*\+\s*row|row\s*\+\s*result);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {

}
--fcc-editable-region--

console.log(result);
```
