---
id: 660f53ad3d39175c5d4335ac
title: Paso 117
challengeType: 1
dashedName: step-117
---

# --description--

¡Tu generador de pirámides está ahora en un estado terminado, con más funcionalidad de la que planeaste originalmente! El siguiente paso es limpiar tu código.

Elimina todos los comentarios, tanto de una sola línea como de múltiples líneas.

# --hints--

No Debes tener comentarios de una sola línea en tu código.

```js
assert.notMatch(code, /\/\//);
```

No Debes tener comentarios de varias líneas en tu código.

```js
assert.notMatch(code, /(?:\*\/|\/\*)/);
```

No Debes tener ningún comentario en tu código.

```js
assert.equal(code, __helpers.removeJSComments(code));
```

Debes eliminar el código que fue comentado por comentarios multilíneas.

```js
assert.notMatch(code, /while/);
assert.lengthOf(code.match(/for\s*\(/g), 2)
assert.lengthOf(code.match(/rows\.push/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
