---
id: 660f48e1d3682f39e81843c4
title: Paso 94
challengeType: 1
dashedName: step-94
---

# --description--

Tu generador de pirámides sigue funcionando. Sin embargo, podría ser posible terminar con un ciclo infinito de nuevo.

Porque solo estás comprobando si `done` no es igual a `count`, si `done` fuera a ser **mayor** que `count` tu ciclo continuaría para siempre.

Actualiza la condición de tu bucle para comprobar si `done` es menor o igual que `count`.

# --hints--

Tu bucle `while` comprobará si `done` es menor o igual que `count`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*done\s*<=\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
