---
id: 660f4efcb8068e4cb470dca1
title: Paso 109
challengeType: 1
dashedName: step-109
---

# --description--

Ahora que has probado estos métodos, puedes hacer otro enfoque de pirámide invertida. Pero primero tienes que limpiar tus experimentos.

Elimina tu array `numbers` y las llamadas a métodos y console.log's.

# --hints--

No debes tener un array `numbers`.

```js
assert.notMatch(__helpers.removeJSComments(code), /numbers/);
```

No deberías tener una variable `unshifted`.

```js
assert.notMatch(__helpers.removeJSComments(code), /unshifted/);
```

No deberías tener la variable `shifted`.

```js
assert.notMatch(__helpers.removeJSComments(code), /shifted/);
```

No debes tener tus declaraciones `console.log`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];
const shifted = numbers.shift();
console.log(shifted);
const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
