---
id: 660f4e74f7fd3f4a99ac2e50
title: Paso 108
challengeType: 1
dashedName: step-108
---

# --description--

Los arreglos también tienen un método `.shift()`. Esto eliminará el **primer** elemento del array, a diferencia de `.pop()` que elimina el último elemento. Aquí hay un ejemplo del método `.shift()`:

```js
const numbers = [1, 2, 3];
numbers.shift();
```

El array `numbers` sería `[2, 3]`.

Justo debajo de tu array `numbers` declara una variable `shifted` y asígnale el resultado de llamar `.shift()` en el array `numbers`. En la siguiente línea, imprime en consola la variable `shifted`.

# --hints--

Debes usar `const` para declarar la variable `shifted`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+shifted/);
```

Deberías llamar a `.shift()` en tu array `numbers`.

```js
assert.deepEqual(numbers, [5, 2, 3]);
```

Debes asignar el resultado de tu llamada `.shift()` a tu variable `shifted`.

```js
assert.equal(shifted, 1);
```

Debes imprimir en consola tu variable `shifted`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*shifted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];


const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
