---
id: 660f2a70ad6225fa503e71c3
title: Paso 46
challengeType: 1
dashedName: step-46
---

# --description--

¡Estás cada vez más cerca! En este punto, te encuentras con lo que se conoce como un <dfn>error por uno</dfn>, un problema frecuente en lenguajes de indexación basados en cero, como JavaScript.

El primer índice de tu array `rows` es `0`, que es la razón por la que inicias tu bucle `for` con `i = 0`. Pero repetir una cadena cero veces da como resultado nada que imprimir.

Para solucionar esto, suma `1` al valor de `i` en tu método `.repeat()`. No lo asignes de vuelta a `i` como lo hiciste en las condiciones del bucle.

# --hints--

Deberías añadir `1` a `i` en tu método `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
