---
id: 660f447efc0e722f016c1be0
title: Paso 81
challengeType: 1
dashedName: step-81
---

# --description--

¡El texto ha vuelto a aparecer! Esto es porque `"false"` es una cadena de texto, que cuando se evalúa a un booleano se convierte en `true`. Esto significa `"false"` es un valor verdadero.

Un valor <dfn>verdadero</dfn> es un valor que se considera verdadero cuando se evalúa como un booleano. La mayoría de los valores que se encuentran en JavaScript serán verdaderos.

Un valor <dfn>falso</dfn> es lo contrario - un valor es considerado falso cuando se evalúa como un booleano. JavaScript tiene una lista definida de valores falsos. Algunos de ellos incluyen `false`, `0`, `""`, `null`, `undefined`, y `NaN`.

Intenta cambiar tu condición `if` a una cadena de texto vacía `""`, que es un valor falso.

# --hints--

Tu condición `if` debe tener `""` como condición.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*("|')\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("false") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
