---
id: 660f0ee51d7460ce88cd248d
title: Paso 29
challengeType: 1
dashedName: step-29
---

# --description--

¿Estabas esperando ver `4` en la consola? El método `.push()` devuelve la nueva longitud del array, después de añadir el valor que le da.

Es importante tener en cuenta qué valores devuelve un método. Toma algo de tiempo para experimentar con los métodos `.push()` y `.pop()`. Cuando estés listo, elimina todos tus métodos `.push()` y `.pop()`, y tus `console.log`.

# --hints--

No deberías tener ningún método `.push()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.push\(/);
```

No deberías tener ningún método `.pop()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.pop\(/);
```

No deberías tener ningún console. log.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log/);
```

No deberías tener una variable `popped`.

```js
assert.notMatch(__helpers.removeJSComments(code), /popped/);
```

No deberías tener una variable `pushed`.

```js
assert.notMatch(__helpers.removeJSComments(code), /pushed/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let pushed = rows.push("freeCodeCamp");
console.log(pushed);
let popped = rows.pop();
console.log(popped);
console.log(rows);
--fcc-editable-region--
```
