---
id: 660f3b664421471aa595170f
title: Paso 74
challengeType: 1
dashedName: step-74
---

# --description--

La pirámide ahora se ve un poco divertida. Porque estás iniciando el bucle en `1` en lugar de `0`, no necesitas añadir uno a `i` cuando lo pasas a `padRow`.

Actualiza el primer argumento de tu llamada `padRow` para que sea `i`.

# --hints--

No debes pasar `i + 1` a `padRow`.

```js
assert.notMatch(__helpers.removeJSComments(code), /padRow\(\s*i\s*\+\s*1/);
```

Deberías pasar `i` a `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*i\s*,/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
