---
id: 660f39b444fd6f16d1e49c1f
title: Paso 73
challengeType: 1
dashedName: step-73
---

# --description--

En lugar de tener que pasar `i + 1` a tu llamada `padRow`, podrías empezar tu bucle en `1`. Esto te permitirá crear un bucle de índice uno.

Actualiza tu iterador para empezar en `1` en lugar de `0`.

# --hints--

Tu bucle `for` debería inicializar `i` en `1`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
