---
id: 660f4de78f775e480ba2e451
title: Paso 107
challengeType: 1
dashedName: step-107
---

# --description--

El método `.unshift()` de un array te permite agregar un valor al **comienzo** del array, a diferencia de `.push()` que añade el valor al final del array. `.unshift()` retorna la nueva longitud del arreglo desde el cual fue llamado.

```js
const countDown = [2, 1, 0];
const newLength = countDown.unshift(3);
console.log(countDown); // [3, 2, 1, 0]
console.log(newLength); // 4
```

Utiliza `const` para declarar la variable `unshifted` y asígnale el resultado de llamar a `.unshift()` en tu array `numbers`. Pasa `5` como argumento. Luego imprime tu variable `unshifted`.

# --hints--

Debes declarar tu variable `unshifted`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+unshifted/);
```

Debes llamar al método `.unshift()` en tu array `numbers`.

```js
assert.match(__helpers.removeJSComments(code), /numbers\.unshift\(/);
```

Debes pasar `5` como argumento a tu llamada `.unshift()`.

```js
assert.deepEqual(numbers, [5,1,2,3]);
```

Debes asignar el resultado de tu llamada `.unshift()` a tu variable `unshifted`.

```js
assert.equal(unshifted, 4);
```

Debes imprimir tu variable `unshifted`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*unshifted\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];

console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
