---
id: 660f4f79e2a82a4e92290f44
title: Paso 110
challengeType: 1
dashedName: step-110
---

# --description--

A veces deseas traer el código anterior que has comentado. Puedes hacerlo eliminando la `/*` y la `*/` alrededor de ese código. Esto se llama <dfn>descomentar</dfn>.

Descomenta solo tu primer bucle `for`. Deja el comentario de una sola línea y los otros dos comentarios de varias líneas en su lugar.

# --hints--

Deberías descomentar tu primer bucle `for`.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /for\s*\(\s*let\s+i\s*=\s*1;\s*i\s*<=\s*count;\s*i\+\+\s*\)/)
```

No deberías eliminar el comentario de una sola línea.

```js
assert.match(code, /\/\/\sTODO:/);
```

No debes descomentar tu bucle `while`.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

No debes descomentar tu segundo bucle `for`.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for\s*\(/g), 2)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
