---
id: 660f033cf051ebb50ea3bf48
title: Paso 18
challengeType: 1
dashedName: step-18
---

# --description--

En programación, a menudo necesitarás trabajar con muchos datos. Hay muchas estructuras de datos que pueden ayudarte a organizar y administrar los datos. Una de las estructuras de datos más básicas es un <dfn>array</dfn>.

Un <dfn>array</dfn> es un tipo de datos no primitivo que puede contener una serie de valores. Los tipos de datos no primitivos difieren de los tipos de datos primitivos en que pueden contener datos más complejos. Tipos de datos primitivos como cadenas de texto y números sólo pueden contener un valor a la vez.

Los arrays se denotan usando corchetes (`[]`). Este es un ejemplo de una variable con el valor de un array vacío:

```js
let array = [];
```

Declarar una variable `rows` y asígnale un array vacío.

# --hints--

Deberías tener una variable `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows/);
```

Deberías usar `let` para declarar la variable `rows`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+rows/);
```

Deberías asignar un array vacío a tu variable `rows`.

```js
assert.deepEqual(rows, []);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
