---
id: 660f4377a359972c521d3f4b
title: Paso 79
challengeType: 1
dashedName: step-79
---

# --description--

Verás la cadena de texto impresa en la consola, porque `true` es de hecho verdadero.

Cambia la condición de tu declaración `if` al valor booleano `false`.

# --hints--

Tu condición `if` debe tener `false` como condición.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if (true) {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
