---
id: 660f07d231941bc11719f664
title: Paso 21
challengeType: 1
dashedName: step-21
---

# --description--

Los arrays son especiales en el sentido de que se consideran <dfn>mutables</dfn>. Esto significa que puedes cambiar el valor directamente en un índice.

Por ejemplo, este código asignaría el número `25` al segundo elemento del array:

```js
let array = [1, 2, 3];
array[1] = 25;
console.log(array); // prints [1, 25, 3]
```

Actualiza el **tercer** elemento de tu array `rows` para que pase a ser el número `10`. Luego imprime el array `rows` en tu consola.

# --hints--

Deberías usar nuevamente la notación de corchetes en tu array `rows`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows\[/g), 2)
```

Debes acceder al tercer elemento del array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]/);
```

Debes usar el operador de asignación en el tercer elemento del array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]\s*=/);
```

Debes asignar el valor `10` al tercer elemento de tu array `rows`.

```js
assert.equal(rows[2], 10);
```

Deberías tener un segundo `console.log` en tu código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Deberías imprimir el array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--

--fcc-editable-region--
```
