---
id: 660f165270622fd4ec0da3f7
title: Paso 31
challengeType: 1
dashedName: step-31
---

# --description--

Declarar una variable con la palabra clave `let` permite reasignarla. Esto significa que podrías cambiar `character` por un valor completamente diferente.

Para este proyecto, no querrás cambiar los valores de estas variables. Entonces, en su lugar, deberías usar `const` para declararlos. Las constantes `const` son especiales.

Primero, una constante `const` no se puede reasignar como una variable `let`. Este código arrojaría un error:

```js
const firstName = "Naomi";
firstName = "Jessica";
```

Una constante `const` tampoco se puede declarar sin inicializar. Este código arrojaría un error:

```js
const firstName;
```

Reemplaza tus palabras reservadas `let` con `const`.

# --hints--

Deberías usar `const` para declarar la constante `character`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+character/);
```

Deberías usar `const` para declarar la constante `count`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+count/);
```

Deberías usar `const` para declarar la constante `rows`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+rows/);
```

No deberías usar `let` en tu código.

```js
assert.notMatch(__helpers.removeJSComments(code), /let/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
let count = 8;
let rows = [];
--fcc-editable-region--
```
