---
id: 660f4b641290da41b2cf0dd9
title: Paso 103
challengeType: 1
dashedName: step-103
---

# --description--

Al igual que las adiciones, hay diferentes operadores que se pueden utilizar para la sustracción. El operador de <dfn>asignación de sustracción</dfn> `-=` resta el valor dado del valor de la variable actual, entonces asigna el resultado de vuelta a la variable.

Reemplaza la instrucción del iterador con la sentencia correcta usando el operador de asignación de sustracción.

# --hints--

Tu bucle `for` no debe usar `i = i - 1`.

```js
assert.notMatch(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*i\s*=\s*i\s*-\s*1\s*\)\s*\{\s*rows\.push\(\s*padRow\s*\(\s*i\s*,\s*count\s*\)\s*\);/);
```

Tu bucle `for` debería usar la asignación de sustracción para reducir `i` por `1`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*i\s*-=\s*1\s*\)\s*\{\s*rows\.push\(\s*padRow\s*\(\s*i\s*,\s*count\s*\)\s*\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; i > 0; i = i - 1) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
