---
id: 660f2b6fd54ac1fc142804dd
title: Paso 47
challengeType: 1
dashedName: step-47
---

# --description--

La lógica para formatear esta pirámide probablemente se complicará, lo que significa que es un buen momento para extraer ese código en una función.

Una <dfn>función</dfn> es un bloque de código que puede reutilizarse a lo largo de tu aplicación. Las funciones se declaran con la siguiente sintaxis:

```js
function name(parameter) {

}
```

La palabra reservada `function` le dice a JavaScript que la variable `name` va a ser una función. `parameter` es una variable que representa un valor que se pasa a la función cuando se utiliza. Una función puede tener tantos <dfn>parámetros</dfn> como queramos. Como un bucle `for`, el espacio entre llaves es el <dfn>cuerpo de la función</dfn>.

Declara una función `padRow`. No crees ninguna variable como parámetro todavía. El cuerpo de la función debe estar vacío. Recuerda que debes utilizar la notación camelCase para nombrar tu función.

# --hints--

Debes usar la palabra reservada `function`.

```js
assert.match(__helpers.removeJSComments(code), /function/);
```

Debes declarar una función `padRow`.

```js
assert.isFunction(padRow);
```

Tu función `padRow()` no debería tener ningún parámetro.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)/);
```

Tu función `padRow()` debe tener el cuerpo vacío.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
