---
id: 660f280dda5040f707c76b4a
title: Paso 45
challengeType: 1
dashedName: step-45
---

# --description--

Ahora tienes una serie de caracteres `#`, pero aún falta la forma de la pirámide. Afortunadamente, la variable `i` representa el número actual de la "fila" en tu bucle, permitiéndote usarlo para crear una estructura similar a la pirámide.

Para lograr esto, se utilizará el método `.repeat()` disponible para cadenas de texto. Este método acepta un número como argumento, especificando el número de veces para repetir la cadena de texto de destino. Por ejemplo, usando `.repeat()` para generar la cadena de texto `"Code! Code! Code!"`:

```js
const activity = "Code! ";
activity.repeat(3);
```

Usa el método `.repeat()` en tu constante `character` y dale el valor de `i` para el número de repeticiones.

# --hints--

Debes utilizar el método `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /\.repeat\(/);
```

Deberías utilizar el método `.repeat()` en tu constante `character`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(/);
```

Deberías pasar `i` como argumento a tu método `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\)/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
