---
id: 660f3dd626be3a1ffe27e5d1
title: Paso 77
challengeType: 1
dashedName: step-77
---

# --description--

JavaScript también tiene soporte para comentarios de múltiples líneas. Un comentario multi-línea comienza con `/*` y termina con `*/`.

A diferencia de un comentario de una sola línea, un comentario de varias líneas encapsulará múltiples líneas.

Usa `/*` y `*/` para convertir tu bucle `for`, incluyendo el cuerpo, en un comentario multilínea.

# --hints--

Debes iniciar un comentario multi-línea con `/*`.

```js
assert.match(code, /\/\*/);
```

Debes terminar el comentario multi-línea con `*/`.

```js
assert.match(code, /\*\//);
```

Todo tu bucle `for` debe ser comentado.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for/g), 1);
assert.notMatch(stripped, /rows\.push/);
assert.notMatch(stripped, /i <= count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
