---
id: 660f3ce51f70571e1c5227c8
title: Paso 76
challengeType: 1
dashedName: step-76
---

# --description--

Los comentarios pueden ser útiles para explicar por qué tu código toma un enfoque determinado o dejar notas para tu yo del futuro.

En JavaScript, puedes usar `//` para dejar un comentario de una sola línea en tu código.

Add a single-line comment above your `for` loop to remind yourself to change the code to a different kind of loop.

# --hints--

Deberías iniciar un comentario de una sola línea con `//`.

```js
assert.match(code, /\/\//);
```

Tu comentario de una sola línea debe tener al menos cinco caracteres.

```js
assert.match(code, /\/\/.{5,}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--

for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
