---
id: 660f4934fb48f63abd5ae371
title: Paso 95
challengeType: 1
dashedName: step-95
---

# --description--

Usar `done` para rastrear el número de filas que se han generado es funcional, pero en realidad se puede limpiar la lógica un poco más.

Los arrays tienen una propiedad especial `length` que te permite ver cuántos valores, o <dfn>elementos</dfn>, hay en el array. Accederías a esta propiedad usando sintaxis como `myArray.length`.

Note that `rows.length` in the `padRow` call would give you an off-by-one error, because `done` is incremented *before* the call.

Update your condition to check if `rows.length` is less than `count`.

# --hints--

Your `while` loop should check if `rows.length` is less than `count`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*rows\.length\s*<\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done <= count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
