---
id: 660f38c34a4de6141c0c369f
title: Paso 71
challengeType: 1
dashedName: step-71
---

# --description--

Tu generador de pirámides ahora funciona como se esperaba. Pero esta es una excelente oportunidad para seguir explorando el código que has escrito.

El operador de adición no es la única manera de agregar valores a una variable. El operador de <dfn>asignación</dfn> puede utilizarse como abreviación para "tomar el valor original de la variable, añadir este valor y asignar el resultado de vuelta a la variable." Por ejemplo, estas dos declaraciones producirían el mismo resultado:

```js
test = test + 1;
test += 1;
```

Actualice su sentencia de iteración en el bucle `for` para usar el operador de asignación de adición.

# --hints--

Tu bucle `for` no debe usar `i = i + 1`;

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

Tu bucle `for` debería usar el operador de adición con `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
