---
id: 660f1e3f047bf4e403268713
title: Paso 39
challengeType: 1
dashedName: step-39
---

# --description--

Desafortunadamente, ahora no puedes ver lo que está haciendo tu bucle.

Usa `let` para declarar una variable `result` y asígnale una cadena de texto vacía. Una cadena de texto vacía es representada por comillas sin nada entre ellas, como `""`.

# --hints--

Debes declarar una variable `result`.

```js
assert.match(__helpers.removeJSComments(code), /result/);
```

Deberías usar `let` para declarar tu variable `result`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+result/);
```

Tu variable `result` debe ser una cadena de texto vacía.

```js
assert.equal(result, "");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

--fcc-editable-region--

--fcc-editable-region--
```
