---
id: 660f48a419b40238e2b8b4d5
title: Paso 93
challengeType: 1
dashedName: step-93
---

# --description--

Tu bucle ya no depende de la variable `continueLoop`. Esto hace a la variable una declaración <dfn>no utilizada</dfn>. Por lo general, quierrás evitar declaraciones no utilizadas para evitar confusiones futuras.

Elimina tu variable `continueLoop`.

# --hints--

Ya no deberías tener una variable `continueLoop`.

```js
assert.notMatch(__helpers.removeJSComments(code), /continueLoop/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
