---
id: 660f4a1472f8e63d76162ce5
title: Paso 99
challengeType: 1
dashedName: step-99
---

# --description--

¿Qué pasaría si hicieras tu piramida hacia abajo, o <dfn>invertida</dfn>? ¡Es hora de probarlo!

Empieza por crear un nuevo bucle `for`. Declara tu iterador `i` y asígale el valor de `count`, luego use el valor booleano `false` para la declaración de la condición e iteración.

# --hints--

Tu código debe tener un bucle `for`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for\s*\(/g), 2);
```

Tu bucle `for` debería inicializar `i` con el valor de `count`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count/);
```

Tu bucle `for` debería usar `false` como condición.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*false/);
```


Tu bucle `for` debería usar `false` como iteración.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*false\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
