---
id: 660f23b53db70af0f2620e78
title: Paso 43
challengeType: 1
dashedName: step-43
---

# --description--

Ahora todos tus números aparecen en la misma línea. Esto no funcionará para la creación de la pirámide.

Necesitarás añadir una nueva línea a cada fila. Sin embargo, pulsando la tecla enter para insertar un salto de línea entre comillas en JavaScript resultará en un error de análisis. En su lugar, necesitas usar la <dfn>secuencia de escape</dfn> especial `\n`, que se interpreta como una nueva línea cuando se imprima en consola la cadena de texto. Por ejemplo:

```js
lineOne = lineOne + "\n" + lineTwo;
```

Use a second addition operator to append a new line after the existing `result` value and the added `row` value.

# --hints--

Debes usar la secuencia de escape `\n`. Recuerda que tiene que ser una cadena de texto, así que está envuelta en comillas.

```js
assert.match(__helpers.removeJSComments(code), /('|")\\n\1/);
```

You should concatenate your `row` variable to your `result` variable.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row\s*\+\s*('|")\\n\1/);
```

You should concatenate the `\n` escape sequence to your `row` variable.

```js
assert.match(__helpers.removeJSComments(code), /row\s*\+\s*('|")\\n\1/);
```

Deberías asignar toda la concatenación a tu variable `result`. No olvides añadir el punto y coma.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*result\s*\+\s*row\s*\+\s*('|")\\n\1;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
