---
id: 660f0da9bf1035c9097af20a
title: Paso 28
challengeType: 1
dashedName: step-28
---

# --description--

Debiste ver `"freeCodeCamp"` impreso en la consola. Esto es porque el método `.pop()` devuelve el valor que fue eliminado del array - y tú *"pusheaste"* `"freeCodeCamp"` al final del array anteriormente.

Pero ¿Qué devuelve el método `.push()`? Asigna tu `rows.push()` a una nueva variable `pushed` y muéstrala en consola.

# --hints--

Debes declarar una variable `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /pushed/);
```

Deberías usar `let` para declarar la variable `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+pushed/);
```

Deberías asignar `rows.push("freeCodeCamp")` a tu variable `pushed`.

```js
assert.equal(pushed, 4);
```

Deberías imprimir en consola tu variable `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*pushed\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
rows.push("freeCodeCamp");

--fcc-editable-region--
let popped = rows.pop();
console.log(popped);
console.log(rows);
```
