---
id: 660f44f10ea40f300b896a5e
title: Paso 84
challengeType: 1
dashedName: step-84
---

# --description--

Ahora que has practicado trabajar con las condiciones `if...else if...else`, puedes eliminarlas de tu código.

Una vez que completes eso, usa `let` para declarar una variable `continueLoop` y asignarle el valor booleano `false`. Luego usa `let` para declarar una variable `done` y asignarle el valor `0`.

# --hints--

No debes tener una condición `if`.

```js
assert.notMatch(__helpers.removeJSComments(code), /if\s*\(/);
```

No debes tener una condición `else if`.

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

No debes tener una condición `else`.

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s*\{/);
```

Deberías usar `let` para declarar una variable `continueLoop`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+continueLoop/);
```

Tu variable `continueLoop` debe tener el valor `false`.

```js
assert.isFalse(continueLoop);
```

Deberías usar `let` para declarar una variable `done`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+done/);
```

Tu variable `done` debe tener el valor `0`.

```js
assert.strictEqual(done, 0);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
} else {
  console.log("This is the else block");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
