---
id: 660f20473aef47e9b8c9afc6
title: Paso 40
challengeType: 1
dashedName: step-40
---

# --description--

Añade una sentencia log para imprimir el valor de `result`. Dependiendo de la consola que utilices, es posible que no veas nada impreso.

# --hints--

Debes añadir un `console.log`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

Deberías imprimir en consola tu variable `result`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--
```
