---
id: 660f18f059fe0fda192ce394
title: Paso 34
challengeType: 1
dashedName: step-34
---

# --description--

Ahora tu bucle necesita un iterador adecuado. El <dfn>iterator</dfn> es una variable que se puede declarar específicamente en tu bucle `for` para controlar cómo el bucle itera o pasa por su lógica.

Es una convención común utilizar `i` como variable iteradora en un bucle. Un bucle `for` te permite declarar esto entre paréntesis `()`. Por ejemplo, aquí hay un bucle `for` que declara una variable `index` y le asigna el valor `100`.

```js
for (let index = 100; "second"; "third") {

}
```

Reemplaza la cadena de texto `"iterator"` con una declaración `let` para la variable `i`. Asígnale el valor `0` para comenzar. Esto le dará a la variable `i` el valor `0` la **primera vez** que se ejecuta el bucle.

# --hints--

Deberías usar `let` para declarar la variable `i`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i/);
```

Deberías asignarle `0` a tu variable `i`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i\s*=\s*0/);
```

Tu bucle `for` debería iniciar un iterador `i` en `0`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for ("iterator"; "condition"; "iteration") {

}
--fcc-editable-region--
```
