---
id: 660f34e99571070d56d2f231
title: Paso 67
challengeType: 1
dashedName: step-67
---

# --description--

Tu función `padRow` tiene dos parámetros que has definido. Los valores se proporcionan a esos parámetros cuando se llama a esa función.

Los valores que proporcionas a una llamada de función se denominan <dfn>argumentos</dfn>, y <dfn>pasas</dfn> argumentos a una llamada de función. Aquí hay una llamada de función con `"Hello"` pasado como argumento:

```js
test("Hello");
```

Pasa `i + 1` y `count` como los argumentos para su llamada `padRow`. Como los parámetros, los argumentos también están separados por una coma.

# --hints--

Deberías pasar `i + 1` a tu llamada `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1/);
```

Debes tener una coma después de tu argumento `i + 1`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*/);
```

Debe pasar `count` como tu segundo argumento.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*count\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow())
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
