---
id: 660f1cedf3676fe26122ebf6
title: Paso 38
challengeType: 1
dashedName: step-38
---

# --description--

Deberías ver los números del cero al siete impresos en tu consola, uno por línea. Esto servirá como base para generar tu pirámide.

Reemplaza tu console. log con una declaración de un push a tu arreglo `rows` para enviar el valor de `i`.

# --hints--

No deberías tener un `console.log`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Deberías usar el método `.push()` en tu array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

Deberías "pushear" el valor de `i` a tu array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

Tu `.push()` debería ejecutarse en tu bucle `for`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*rows\.push\(\s*i\s*\);?\s*\}/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  console.log(i);
}
--fcc-editable-region--
```
