---
id: 660f3ba3cceef11b6ba08b59
title: Paso 75
challengeType: 1
dashedName: step-75
---

# --description--

Desafortunadamente, ahora ha desaparecido la parte inferior de la pirámide. Esto se debe a que has creado otro <dfn>error de partida</dfn>.

Tu bucle original fue para valores `i` de `0` a `7`, porque `count` es `8` y tu condición requiere que `i` sea menor que `count`. Tu bucle se está ejecutando ahora para valores `i` de `1` a `7`.

Tu bucle necesita ser actualizado para ejecutarse cuando `i` también sea `8`. Mirando tu lógica, esto significa que tu bucle debería ejecutarse cuando `i` es <dfn>menor o igual que</dfn> `count`. Puedes usar el operador menor o igual que `<=` para esto.

Actualiza la condición de tu bucle para que se ejecute mientras `i` es menor o igual que `count`.

# --hints--

Tu bucle `for` no debería comprobar si `i` es menor que `count`.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

Tu bucle `for` debe comprobar si `i` es menor o igual que `count`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<=\s*count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
