---
id: 6610b9f7619764fad5fd516d
title: Paso 11
challengeType: 1
dashedName: step-11
---

# --description--

También puedes asignar el valor de una variable a otra variable. Por ejemplo:

```js
let first = "One";
let second = "Two";
second = first;
```

La variable `second` ahora tendría el valor `"One"`.

Para ver esto en acción, cambia la asignación de tu variable `secondCharacter` de `"Test"` a tu variable `character`.

Luego abre la consola para ver lo que se imprime.

# --hints--

No deberías asignar el valor `"Test"` a tu variable `secondCharacter`.

```js
assert.notEqual(secondCharacter, "Test");
```

Deberías asignar el valor de la variable `character` a tu variable `secondCharacter`. No olvides el punto y coma.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=\s*character;/);
```

Tu variable `secondCharacter` ahora debería tener el valor `"World"`.


```js
assert.equal(secondCharacter, "World");
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
--fcc-editable-region--
secondCharacter = "Test";
--fcc-editable-region--
console.log(secondCharacter);
```
