---
id: 6610c424b7119919b62932f4
title: Paso 53
challengeType: 1
dashedName: step-53
---

# --description--

Una función no tiene que devolver un valor "harcodeado". Puede devolver el valor almacenado en una variable. Los parámetros son variables especiales para una función, por lo que también se pueden retornar.

Cambia tu función `padRow` para que haga el `return` del parámetro `name`.

# --hints--

Tu función `padRow` debe devolver el valor del parámetro `name`.

```js
assert.equal(padRow("Naomi"), "Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
