---
id: 6610c77d50636722e5b6be17
title: Paso 60
challengeType: 1
dashedName: step-60
---

# --description--

Debajo de la instrucción `return`, muestra en consola la cadena de texto `"This works!"`.

Después de hacer esto, verás que la cadena `"This works!"` no se muestra en la consola y la línea del `console.log("This works!")` aparece en color gris.

Copia el console.log y pégalo por encima de la instrucción `return`. Ahora, la cadena de texto `"This works!"` debería aparecer en la consola.

Una cosa importante que hay que saber sobre la palabra reservada `return` es que no solo define un valor para ser devuelto desde tu función, también detiene la ejecución de la función. Esto significa que cualquier código después de un `return` no se ejecutará.

# --hints--

Your `padRow` function should return the `test` variable.

```js
assert.equal(padRow("Naomi"), "Testing");
```

Tu primer `console.log` debería venir después de tu palabra reservada `return`.

```js
assert.match(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

Tu segundo `console.log` debería venir antes de tu palabra reservada `return`.

```js
assert.match(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";

  return test;

}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
