---
id: 6610c6541c82551f95e765ab
title: Paso 58
challengeType: 1
dashedName: step-58
---

# --description--

Las variables también pueden ser declaradas dentro de una función. Estas variables se consideran en el ámbito (scope) <dfn>local</dfn> o <dfn>ámbito de bloque</dfn>. Una variable declarada dentro de una función sólo puede ser usada dentro de esa función. Si intentas acceder a ella fuera de la función, obtendrás un error de referencia.

Para ver esto en acción, usa `const` para declarar una variable `test` en tu función `padRow`. Inicialízala con el valor `"Testing"`.

Luego, debajo de tu función, intenta hacer un console.log de `test`. Verás un error porque no está definido fuera del ámbito local de la función. Elimina ese `console.log` para pasar las pruebas y continuar.

# --hints--

Tu función debe declarar una variable `test`.

```js
assert.match(padRow.toString(), /var\s+test/);
```

Debes inicializar la variable `test` con el valor `"Testing"`. No olvides añadir el punto y coma.

```js
assert.match(padRow.toString(), /var\s+test\s*=\s*('|")Testing\1;/)
```

Tu variable `test` debe aparecer antes de la palabra reservada `return`.

```js
const str = padRow.toString();
const testIndex = str.indexOf("test");
const returnLocation = str.indexOf("return");
const testLocation = testIndex !== -1 ? testIndex : returnLocation
assert.isBelow(testLocation, returnLocation);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {

  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
