---
id: 6610c71600966a2191d3a64a
title: Paso 59
challengeType: 1
dashedName: step-59
---

# --description--

Los valores devueltos de una función se utilizan al llamar a la función. Puedes usar la llamada a la función directamente como el valor que regresa, o capturar el valor devuelto en una variable. De esta manera, puedes utilizar el valor asignado a una variable de ámbito local, fuera de la función en la que fue creada.

```js
function getName() {
  const name = "Camper cat";
  return name;
}

console.log(getName()); // "Camper cat"

const capturedReturnValue = getName();
console.log(capturedReturnValue); // "Camper cat"

console.log(name); // reference error
```

Para traer el valor `"Testing"` de la función `padRow` al ámbito global, actualiza la instrucción `return` para devolver sólo la variable `test`.


# --hints--

Tu función `padRow` debería devolver la variable `test`.

```js
assert.equal(padRow("Naomi"), "Testing");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";  
  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
