---
id: 6610c105bbdacc114d6cdc44
title: Paso 48
challengeType: 1
dashedName: step-48
---

# --description--

Para poder usar una función, necesitas llamarla. Una <dfn>llamada a la función</dfn> le dice a tu aplicación que ejecute el código desde dondequiera que elijas llamar a la función. La sintaxis para llamar a una función es el nombre de función seguido de paréntesis. Por ejemplo, este código define y llama a una función `test`.

```js
function test() {

}
test();
```

Llama a tu función `padRow`.

# --hints--

Debes llamar a la función `padRow`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

El cuerpo de la función `padRow` debe estar vacío.

```js
assert.match(padRow.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
