---
id: 6610bf6fa14d700beed1b109
title: Paso 87
challengeType: 1
dashedName: step-87
---

# --description--

El operador de <dfn>igualdad</dfn> `==` se utiliza para comprobar si dos valores son iguales. Para comparar dos valores, usarías una proposición como un `value == 8`.

Below `done++` inside your loop, add an `if` statement. La sentencia debe verificar si `done` es igual a `count` usando el operador de igualdad.

# --hints--

You should use an `if` statement in your loop. It should be added after `done++`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if/);
```

Tu sentencia `if` debe usar el operador de igualdad para comparar `done` y `count` en la condición.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*==\s*count|count\s*==\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
