---
id: 6610c1d97b1671140f95cfbb
title: Paso 50
challengeType: 1
dashedName: step-50
---

# --description--

Ahora añade un console.log para imprimir el valor de tu variable `call`.

# --hints--

Deberías añadir un `console.log`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2)
```

Deberías imprimir en consola tu variable `call`. No olvides añadir el punto y coma.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*call\s*\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
