---
id: 6610c16c4fa0df12c0e30675
title: Paso 49
challengeType: 1
dashedName: step-49
---

# --description--

Estás llamando a tu función `padRow`, pero no estás haciendo nada con esa llamada de función. Todas las funciones en JavaScript <dfn>devuelven</dfn> un valor, lo que significa que proporcionan el resultado definido de llamarlas para que las uses en otros lugares.

Para ver el resultado de llamar a tu función `padRow`, declara una variable `call` y asígnale la llamada a la función `padRow` a esa variable.

# --hints--

Debes declarar una variable `call`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+call/);
```

Debes usar `const` para declarar tu variable `call`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call/);
```

Deberías asignar el resultado de tu llamada `padRow` a la variable `call`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call\s*=\s*padRow\(\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow() {

}
--fcc-editable-region--
padRow();
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
