---
id: 6610bbed59bc2a0194d85533
title: Paso 22
challengeType: 1
dashedName: step-22
---

# --description--

¿Ves cómo se ha cambiado el valor dentro de tu array `rows` directamente? A esto se le llama `mutación`. A medida que aprendas más sobre arrays, aprenderás cuándo mutar un array, y cuándo no hacerlo.

Antes de seguir adelante, esta es una gran oportunidad para aprender un uso común de arrays. Actualmente, tu código accede al último elemento del array con `rows[2]`. Pero puede que no sepas cuántos elementos hay en un array cuando necesitas el último valor.

Puedes hacer uso de la propiedad `.length` de un array - esto devuelve el número de elementos en el array. Para obtener el último elemento de cualquier array, puedes utilizar la siguiente sintaxis:

```js
array[array.length - 1]
```

`array.length` devuelve el número de elementos en el array. Al restar `1`, obtienes el índice del último elemento en el array. Puedes aplicar este mismo concepto a tu array `rows`.

Modifica tu array `rows[2]` para acceder dinámicamente al último elemento en el array `rows`. Consulta el ejemplo anterior para ayudarte.

No Debes ver ningún cambio en tu consola.

# --hints--

Debes usar la propiedad `.length` de tu array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.length/);
```

Debes usar `rows.length` en tu notación de corchetes.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length/);
```

Debes restar `1` de la longitud de la notación de corchetes.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--
rows[2] = 10;
--fcc-editable-region--
console.log(rows);
```
