---
id: 6610c8cfe4cf4d278e35c156
title: Paso 63
challengeType: 1
dashedName: step-63
---

# --description--

Como se esperaba, tu función ahora devuelve `undefined` de nuevo. Tu variable `call` ya no es necesaria. así que quita la declaración `call` y el `console.log` para la variable `call`.

# --hints--

No debes tener una declaración `call`.

```js
assert.notMatch(__helpers.removeJSComments(code), /const\s*call/);
```

No deberías imprimir en consola tu variable `call`.

```js
assert.notMatch(__helpers.removeJSComments(code), /call/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
