---
id: 6610c87eac0f0b256d7b037e
title: Paso 62
challengeType: 1
dashedName: step-62
---

# --description--

Debido a que su función ya no utilizaba el parámetro, el cambio de argumento no le afectó.

Adelante, elimina la declaración `test` y la declaración `return` de tu función `padRow` para que la función esté nuevamente vacía.

# --hints--

Tu función `padRow` no debe tener la variable `test`.

```js
assert.notMatch(padRow.toString(), /test/);
```

Tu función `padRow` no debería devolver un valor.

```js
assert.notMatch(padRow.toString(), /return/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  const test = "Testing";
  return test;
}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
