---
id: 6610b8f6a98d25f4d485a94d
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

El valor por defecto de una variable no inicializada es `undefined`. Este es un tipo de datos especial que representa un valor que no tiene una definición todavía.

Todavía puedes asignar un valor a una variable no inicializada. Aquí hay un ejemplo:

```js
let uninitialized;
uninitialized = "assigned";
```

Asigna la cadena de texto `"Test"` a tu variable `secondCharacter` debajo de tu declaración. Abre la consola para ver cómo ha cambiado tu registro.

# --hints--

No debes inicializar `secondCharacter`. Recuerda que la inicialización significa asignar un valor cuando declaras la variable.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+secondCharacter\s*=/);
```

Deberías usar el operador de asignación en `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=/);
```

Deberías asignar la cadena de texto `"Test"` a tu variable `secondCharacter`.

```js
assert.equal(secondCharacter, "Test");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

console.log(secondCharacter);
--fcc-editable-region--
```
