---
id: 66458f0a05df478aa627629e
title: Paso 82
challengeType: 1
dashedName: step-82
---

# --description--

¡El texto ha desaparecido de nuevo! Las cadenas de texto vacías son evaluadas como `false`, haciéndolas un <dfn>valor falso</dfn>. Aprenderás más sobre los valores verdaderos y falsos en futuros proyectos.

Además de la condición `if`, JavaScript también tiene la condición <dfn>else if</dfn>. La condición `else if` te permite validar múltiples condiciones en un único bloque de código.

Aquí está la sintaxis para la condición `else if`:

```js
if (condition1) {
  // code to run if condition1 is true
} else if (condition2) {
  // code to run if condition2 is true
} else if (condition3) {
  // code to run if condition3 is true
} 
```

Si la primera condición es `false`, JavaScript comprobará la siguiente condición en la cadena. Si la segunda condición es `false`, JavaScript comprobará la tercera condición, y así sucesivamente.

Debajo de la declaración de tu `if`, añade un `else if` que verifique si `5` es menor que `10`. Luego, dentro del cuerpo de condición `else if`, muestra en consola la cadena de texto `"5 is less than 10"`.

Mira la consola para ver los resultados.

# --hints--

Debes tener una condición `else if`.

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

Tu `else if` debería comprobar si `5` es menor que `10`.

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(\s*5\s*<\s*10\s*\)/);
```

Debes imprimir en consola la cadena de texto `"5 is less than 10"`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)5\s+is\s+less\s+than\s+10\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

