---
id: 664ed9344ba3495e1f67f8ea
title: Paso 14
challengeType: 1
dashedName: step-14
---

# --description--

Ahora que has repasado la declaración y la inicialización, elimina el código que escribiste para tu repaso.

No elimines tu variable `character`.

# --hints--

No debes tener la variable `profession`.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+profession/);
```

No debes tener la variable `age`.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+age/);
```

No deberías tener un `console.log` para `age`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*age\s*\);/);
```

No deberías tener un `console.log` para `profession`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*profession\s*\);/);
```

Deberías tener tu variable `character`.

```js
assert.strictEqual(character, "Hello");
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--
let profession = "teacher";
let age;
console.log(age);
console.log(profession);
--fcc-editable-region--
```
