---
id: 664ed97c55a99e5ffab759dc
title: Paso 13
challengeType: 1
dashedName: step-13
---

# --description--

Antes de avanzar, deberías tomarte un momento para revisar los conceptos que has aprendido.

Usa la palabra reservada `let` para declarar una variable `profession` y una variable `age`. Inicializa `profession` con la cadena de caracteres`"teacher"` pero no inicialices `age` con ningún valor.

Haz un console.log de ambas variables para ver los resultados en la consola.

# --hints--

Debes declarar una variable `profession`.

```js
assert.match(code, /(?:var|let|const)\s+profession/);
```

Deberías usar `let` para declarar la variable `profession`.

```js
assert.match(code, /let\s+profession/);
```

Debes asignar la cadena de caracteres `"teacher"` a la variable `profession`.

```js
assert.strictEqual(profession, "teacher");
```

Deberías declarar una variable `age`.

```js
assert.match(code, /(?:var|let|const)\s+age/);
```

Deberías usar `let` para declarar la variable `age`.

```js
assert.match(code, /let\s+age/);
```

No debes asignar ningún valor a la variable `age`.

```js
assert.isUndefined(age);
```

Deberías mostrar en consola la variable `age`.

```js
assert.match(code, /console\.log\(\s*age\s*\);?/);
```

Deberías mostrar en consola la variable `profession`.

```js
assert.match(code, /console\.log\(\s*profession\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--

--fcc-editable-region--
```
