---
id: 665732da4815b70bb083915e
title: Paso 24
challengeType: 1
dashedName: step-24
---


# --description--

En los últimos pasos, aprendiste todo acerca de trabajar con arrays. Tómete un momento para revisar lo que has aprendido.

Comienza declarando una variable `cities` e inicialízala como un arreglo de las cadenas de texto `"London"`, `"New York"`, y `"Mumbai"`. Luego muestra esa variable en la consola.

Después de mostrar la variable, cambia el último elemento de `cities` a la cadena de texto `"Mexico City"`, luego muestra en consola la variable `cities` de nuevo.

Cuando se haya hecho correctamente, deberías ver esta salida en la consola.

```js
[ "London", "New York", "Mumbai" ]
[ "London", "New York", "Mexico City" ]
```

# --hints--

Deberías usar `let` para declarar una variable `cities`.

```js
assert.match(code, /let\s+cities/);
```

Deberías asignar una matriz de las cadenas de texto `"London"`, `"New York"`, y `"Mumbai"` a la variable `cities`.

```js
assert.match(code, /let\s+cities\s*=\s*\[\s*("|'|`)London\1\s*,\s*("|'|`)New York\2\s*,\s*("|'|`)Mumbai\3\s*\]/);
```

Deberías usar `console.log()` para mostrar toda la matriz `cities` en la consola.

```js
assert.match(code, /console\.log\(\s*cities\s*\)/);
```

Deberías actualizar el último elemento del array `cities` a la cadena de texto `"Mexico City"`. Recuerda que puedes acceder al último elemento de un array usando `array[array.length - 1]`.

```js
assert.deepEqual(cities, ["London", "New York", "Mexico City"]);
```

Debes tener dos declaraciones `console.log(cities)` en tu código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*cities\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--

--fcc-editable-region--
console.log(rows);
```
