---
id: 5a858944d96184f06fd60d61
title: Tengeneza chombo chako cha kwanza cha gridi cha CSS
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cqwREC4'
forumTopicId: 301129
dashedName: create-your-first-css-grid
---

# --description--

Badilisha kipengele chochote cha HTML kuwa chombo cha gridi kwa kuweka sifa yake ya `display` kuwa `grid`. Hii inakupa uwezo wa kutumia sifa zote nyingine zinazohusiana na CSS Grid.

**Kumbuka:** Katika CSS Grid, kipengele cha mzazi huitwa <dfn>kontena</dfn> na watoto wake huitwa <dfn>vitu</dfn>.

# --instructions--

Badilisha onyesho la div lenye darasa la `container` kuwa `grid`.

# --hints--

Darasa la `container` linapaswa kuwa na sifa ya `display` yenye thamani ya `grid`.

```js
assert.match(code, /.container\s*?{[\s\S]*display\s*?:\s*?grid\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    display: grid;
  }
</style>
```
