---
id: 5a90375238fddaf9a66b5d3b
title: Panga kipengee kwa wima kwa kutumia align-self
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cmzd4fz'
forumTopicId: 301123
dashedName: align-an-item-vertically-using-align-self
---

# --description--

Kama unavyoweza kupanga kipengee kwa usawa, kuna njia ya kupanga kipengee kwa wima pia. Ili kufanya hivi, unatumia sifa ya `align-self` kwenye kipengee. Sifa hii inakubali thamani zote kama `justify-self` kutoka kwenye changamoto iliyopita.

# --instructions--

Panga kipengee chenye darasa la `item3` kwa wima kwenye `end`.

# --hints--

Darasa la `item3` linapaswa kuwa na sifa ya `align-self` yenye thamani ya `end`.

```js
assert.match(code, /.item3\s*?{[\s\S]*align-self\s*?:\s*?end\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }

  .item3 {
    background: PaleTurquoise;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item3 {
    align-self: end;
  }
</style>
```
