---
id: 5a9036d038fddaf9a66b5d32
title: Ongeza safu kwa grid-template-columns
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c7NzDHv'
forumTopicId: 301117
dashedName: add-columns-with-grid-template-columns
---

# --description--

Kuumba kipengele cha gridi pekee hakutakufikisha mbali. Unahitaji pia kufafanua muundo wa gridi. Ili kuongeza safu kwenye gridi, tumia sifa ya `grid-template-columns` kwenye chombo cha gridi kama ilivyoonyeshwa hapa chini:

```css
.container {
  display: grid;
  grid-template-columns: 50px 50px;
}
```

Hii itatoa safu mbili kwenye gridi yako ambazo kila moja ni pana 50px. Idadi ya vigezo vinavyotolewa kwa sifa ya `grid-template-columns` inaonyesha idadi ya safu kwenye gridi, na thamani ya kila kigezo inaonyesha upana wa kila safu.

# --instructions--

Toa safu tatu kwenye chombo cha gridi ambazo kila moja ni pana `100px`.

# --hints--

Darasa la `container` linapaswa kuwa na sifa ya `grid-template-columns` yenye vitengo vitatu vya `100px`.

```js
const templateColumns = new __helpers.CSSHelp(document).getStyle(
  '.container'
)?.gridTemplateColumns;
assert.strictEqual(templateColumns, '100px 100px 100px');
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: 100px 100px 100px;
  }
</style>
```
