---
id: 5a90374338fddaf9a66b5d3a
title: Panga kipande kwa usawa kwa kutumia justify-self
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpKHq'
forumTopicId: 301122
dashedName: align-an-item-horizontally-using-justify-self
---

# --description--

Katika CSS Grid, maudhui ya kila kipande yanapatikana ndani ya kisanduku kinachoitwa <dfn>seli</dfn>. Unaweza kupanga nafasi ya maudhui ndani ya seli yake kwa usawa kwa kutumia sifa ya `justify-self` kwenye kipande cha gridi. Kwa chaguo-msingi, sifa hii ina thamani ya `stretch`, ambayo itafanya maudhui yajaze upana wote wa seli. Sifa hii ya CSS Grid inakubali thamani nyingine pia:

`start`: inapanga maudhui upande wa kushoto wa seli,

`center`: inapanga maudhui katikati ya seli,

`end`: inapanga maudhui upande wa kulia wa seli.

# --instructions--

Tumia sifa ya `justify-self` kupanga kipande chenye darasa la `item2` katikati.

# --hints--

Darasa la `item2` linapaswa kuwa na sifa ya `justify-self` yenye thamani ya `center`.

```js
assert.match(
  code,
  /.item2\s*?{[\s\S]*justify-self\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }

  .item2 {
    background: LightSalmon;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item2 {
    justify-self: center;
  }
</style>
```
