---
id: 5a9036ee38fddaf9a66b5d37
title: Ongeza mapengo kwa haraka zaidi kwa grid-gap
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/ca2qVtv'
forumTopicId: 301118
dashedName: add-gaps-faster-with-grid-gap
---

# --description--

`grid-gap` ni sifa ya ufupisho kwa `grid-row-gap` na `grid-column-gap` kutoka changamoto mbili zilizopita ambayo ni rahisi zaidi kuitumia. Ikiwa `grid-gap` ina thamani moja, itaunda pengo kati ya safu zote na nguzo zote. Hata hivyo, ikiwa kuna thamani mbili, itatumia thamani ya kwanza kuweka pengo kati ya safu na thamani ya pili kwa nguzo.

# --instructions--

Tumia `grid-gap` kuanzisha pengo la `10px` kati ya safu na pengo la `20px` kati ya nguzo.

# --hints--

Darasa la `container` linapaswa kuwa na sifa ya `grid-gap` inayozindua pengo la `10px` kati ya safu na pengo la `20px` kati ya nguzo.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-gap\s*?:\s*?10px\s+?20px\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>
<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-gap: 10px 20px;
  }
</style>
```
