---
id: 5a9036ee38fddaf9a66b5d34
title: Tumia vipimo vya CSS Grid kubadilisha ukubwa wa safu na nguzo
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cvE8phd'
forumTopicId: 301134
dashedName: use-css-grid-units-to-change-the-size-of-columns-and-rows
---

# --description--

Unaweza kutumia vipimo vya kudumu na vipimo vinavyohusiana kama `px` na `em` katika CSS Grid kuainisha ukubwa wa safu na nguzo. Pia unaweza kutumia hizi:

`fr`: huweka nguzo au safu kuwa sehemu ya nafasi iliyopo,

`auto`: huweka nguzo au safu kuwa upana au urefu wa maudhui yake moja kwa moja,

`%`: huweka nguzo au safu kuwa asilimia ya upana wa kontena yake.

Huu ni msimbo unaozalisha matokeo kwenye onyesho la awali:

```css
grid-template-columns: auto 50px 10% 2fr 1fr;
```

Kipande hiki huunda nguzo tano. Nguzo ya kwanza ni pana kama maudhui yake, nguzo ya pili ni 50px, nguzo ya tatu ni asilimia 10 ya kontena yake, na kwa nguzo mbili za mwisho; nafasi iliyobaki hugawanywa katika sehemu tatu, mbili hutolewa kwa nguzo ya nne, na moja kwa nguzo ya tano.

# --instructions--

Tengeneza gridi yenye nguzo tatu zenye upana kama ifuatavyo: 1fr, 100px, na 2fr.

# --hints--

Darasa lako la `container` linapaswa kuwa na sifa ya `grid-template-columns` inayoweka nguzo tatu zenye upana ifuatayo: `1fr`, `100px`, na `2fr`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?1fr\s*?100px\s*?2fr\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: auto 50px 10% 2fr 1fr;

    /* Only change code above this line */
    grid-template-rows: 50px 50px;
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: 1fr 100px 2fr;
  }
</style>
```
