---
id: 5a90376038fddaf9a66b5d3c
title: Linganisha vitu vyote kwa usawa wa usawa kwa kutumia justify-items
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpECn'
forumTopicId: 301120
dashedName: align-all-items-horizontally-using-justify-items
---

# --description--

Wakati mwingine unataka vitu vyote katika gridi yako ya CSS vishirikiane ulinganifu mmoja. Unaweza kutumia sifa ulizojifunza awali na kuvilinganisha kimoja kimoja, au unaweza kuvilinganisha vyote kwa pamoja kwa usawa kwa kutumia `justify-items` kwenye chombo chako cha gridi. Sifa hii inaweza kupokea thamani zote ulizojifunza katika changamoto mbili zilizopita, tofauti ikiwa ni kwamba itasogeza **vitu vyote** katika gridi yetu kwa ulinganifu unaotakiwa.

# --instructions--

Tumia sifa hii kuleta vitu vyote katikati kwa usawa.

# --hints--

Darasa la `container` linapaswa kuwa na sifa ya `justify-items` yenye thamani ya `center`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*justify-items\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    justify-items: center;
  }
</style>
```
