---
id: 5a9036ee38fddaf9a66b5d35
title: Tengeneza Pengo la Safu kwa Kutumia grid-column-gap
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cVZ8vfD'
forumTopicId: 301124
dashedName: create-a-column-gap-using-grid-column-gap
---

# --description--

Hadi sasa katika gridi ulizotengeneza, safu zote zimekuwa zikigusana moja kwa moja. Wakati mwingine unataka kuweka pengo kati ya safu hizo. Ili kuongeza pengo kati ya safu, tumia sifa ya `grid-column-gap` kama ifuatavyo:

```css
grid-column-gap: 10px;
```

Hii inaunda nafasi tupu ya 10px kati ya safu zetu zote.

# --instructions--

Toa pengo la `20px` kati ya safu katika gridi.

# --hints--

Darasa la `container` linapaswa kuwa na sifa ya `grid-column-gap` yenye thamani ya `20px`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-column-gap\s*?:\s*?20px\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    /* Only change code below this line */

   
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-column-gap: 20px;
  }
</style>
```
