---
id: 5a94fe3669fb03452672e45f
title: Punguza Kurudia kwa Kutumia kitendakazi cha repeat
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cQvqyHR'
forumTopicId: 301133
dashedName: reduce-repetition-using-the-repeat-function
---

# --description--

Ulikotumia `grid-template-columns` na `grid-template-rows` kuunda muundo wa gridi, uliingiza thamani kwa kila safu au safu wima uliyotengeneza.

Tuseme unataka gridi yenye safu 100 za urefu sawa. Si rahisi sana kuingiza thamani 100 moja moja. Kwa bahati nzuri, kuna njia bora zaidi - kwa kutumia kitendakazi cha `repeat` kuonyesha mara ngapi unataka safu au safu wima irudiwe, ikifuatiwa na koma na thamani unayotaka kurudia.

Hapa kuna mfano ambao ungeunda gridi ya safu 100, kila safu ikiwa na urefu wa 50px.

```css
grid-template-rows: repeat(100, 50px);
```

Unaweza pia kurudia thamani nyingi kwa kutumia kitendakazi cha repeat na kuingiza kitendakazi hicho kati ya thamani nyingine wakati wa kuunda muundo wa gridi. Hivi ndivyo inavyoonekana:

```css
grid-template-columns: repeat(2, 1fr 50px) 20px;
```

Hii inatafsiriwa kama:

```css
grid-template-columns: 1fr 50px 1fr 50px 20px;
```

**Kumbuka:** `1fr 50px` inarudiwa mara mbili ikifuatiwa na 20px.

# --instructions--

Tumia `repeat` kuondoa kurudia kutoka kwenye sifa ya `grid-template-columns`.

# --hints--

Darasa la `container` linapaswa kuwa na sifa ya `grid-template-columns` ambayo imewekwa kurudia safu 3 zenye upana wa `1fr`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?3\s*?,\s*?1fr\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: 1fr 1fr 1fr;

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(3, 1fr);
  }
</style>
```
