---
id: 5a94fe8569fb03452672e464
title: Tengeneza mitandao ndani ya mitandao
challengeType: 0
forumTopicId: 301128
dashedName: create-grids-within-grids
---

# --description--

Kugeuza kipengele kuwa mtandao huathiri tu tabia za mzazi wake wa moja kwa moja. Kwa hivyo kwa kugeuza mzazi wa moja kwa moja kuwa mtandao, unakuwa na mtandao ndani ya mtandao.

Kwa mfano, kwa kuweka sifa za `display` na `grid-template-columns` za kipengele chenye darasa la `item3`, unaunda mtandao ndani ya mtandao wako.

# --instructions--

Geuza kipengele chenye darasa la `item3` kuwa mtandao wenye safu mbili zenye upana wa `auto` na `1fr` kwa kutumia `display` na `grid-template-columns`.

# --hints--

Darasa la `item3` linapaswa kuwa na sifa ya `grid-template-columns` yenye thamani za `auto` na `1fr`.

```js
assert.match(
  code,
  /.item3\s*?{[\s\S]*grid-template-columns\s*?:\s*?auto\s*?1fr\s*?;[\s\S]*}/gi
);
```

Darasa la `item3` linapaswa kuwa na sifa ya `display` yenye thamani ya `grid`.

```js
assert.match(code, /.item3\s*?{[\s\S]*display\s*?:\s*?grid\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .container {
    font-size: 1.5em;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: auto 1fr;
    grid-template-rows: auto 1fr auto;
    grid-gap: 10px;
    grid-template-areas:
      'advert header'
      'advert content'
      'advert footer';
  }
  .item1 {
    background: LightSkyBlue;
    grid-area: header;
  }

  .item2 {
    background: LightSalmon;
    grid-area: advert;
  }

  .item3 {
    background: PaleTurquoise;
    grid-area: content;
    /* Only change code below this line */


    /* Only change code above this line */
  }

  .item4 {
    background: lightpink;
    grid-area: footer;
  }

  .itemOne {
    background: PaleGreen;
  }

  .itemTwo {
    background: BlanchedAlmond;
  }
</style>

<div class="container">
  <div class="item1">header</div>
  <div class="item2">advert</div>
  <div class="item3">
    <div class="itemOne">paragraph1</div>
    <div class="itemTwo">paragraph2</div>
  </div>
  <div class="item4">footer</div>
</div>
```

# --solutions--

```html
<style>
  .item3 {
    grid-template-columns: auto 1fr;
    display: grid;
  }
</style>
```
