---
id: 5a94fe4469fb03452672e460
title: Punguza Ukubwa wa Kitu Kutumia kitendakazi cha minmax
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cD97RTv'
forumTopicId: 301131
dashedName: limit-item-size-using-the-minmax-function
---

# --description--

Kuna kitendakazi kingine iliyo ndani cha kutumia pamoja na `grid-template-columns` na `grid-template-rows` kinachoitwa `minmax`. Kinatumika kupunguza ukubwa wa vitu wakati chombo cha gridi kinapobadilika ukubwa. Ili kufanya hivi unahitaji kubainisha anuwai ya ukubwa unaokubalika kwa kitu chako. Hapa kuna mfano:

```css
grid-template-columns: 100px minmax(50px, 200px);
```

Katika msimbo hapo juu, `grid-template-columns` imewekwa kuunda safu mbili; ya kwanza ni upana wa 100px, na ya pili ina upana wa kiwango cha chini wa 50px na kiwango cha juu cha 200px.

# --instructions--

Kutumia kitendakazi cha `minmax`, badilisha `1fr` katika kitendakazi cha `repeat` kwa ukubwa wa safu unao kuwa na upana wa kiwango cha chini wa `90px` na upana wa kiwango cha juu wa `1fr`, kisha badilisha ukubwa wa jopo la onyesho kuona athari.

# --hints--

Darasa la `container` linapaswa kuwa na sifa ya `grid-template-columns` ambayo imewekwa kurudia safu 3 zenye upana wa kiwango cha chini wa `90px` na upana wa kiwango cha juu wa `1fr`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?3\s*?,\s*?minmax\s*?\(\s*?90px\s*?,\s*?1fr\s*?\)\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(3, 1fr);

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(3, minmax(90px, 1fr));
  }
</style>
```
