---
id: 5a94fe1369fb03452672e45d
title: Weka vitu katika maeneo ya gridi kwa kutumia sifa ya grid-area
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cRrqmtV'
forumTopicId: 301132
dashedName: place-items-in-grid-areas-using-the-grid-area-property
---

# --description--

Baada ya kuunda kiolezo cha eneo kwa chombo chako cha gridi, kama ilivyoonyeshwa katika changamoto iliyopita, unaweza kuweka kipengee katika eneo lako maalum kwa kurejelea jina ulilolipa. Ili kufanya hivyo, unatumia sifa ya `grid-area` kwenye kipengee kama hii:

```css
.item1 {
  grid-area: header;
}
```

Hii inamruhusu gridi kujua kwamba unataka darasa la vitu la `item1` liwe katika eneo lililopewa jina `header`. Katika kesi hii, kipengee kitachukua safu yote ya juu kwa sababu safu yote hiyo imepewa jina la eneo la `header`.

# --instructions--

Weka kipengee chenye darasa la vitu la `item5` katika eneo la `footer` kwa kutumia sifa ya `grid-area`.

# --hints--

Darasa la vitu la `item5` linapaswa kuwa na sifa ya `grid-area` yenye thamani ya `footer`.

```js
assert.match(
  __helpers.removeCssComments(code),
  /.item5\s*?{[\s\S]*grid-area\s*?:\s*?footer\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }

  .item5 {
    background: PaleGreen;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    grid-template-areas:
      'header header header'
      'advert content content'
      'footer footer footer';
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item5 {
    grid-area: footer;
  }
</style>
```
