---
id: 5a94fe5469fb03452672e461
title: Tengeneza mipangilio inayobadilika kwa kutumia auto-fill
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cmzdycW'
forumTopicId: 301126
dashedName: create-flexible-layouts-using-auto-fill
---

# --description--

Kitendakazi cha repeat kina chaguo kinachoitwa <dfn>auto-fill</dfn>. Hii inakuwezesha kuingiza moja kwa moja safu au nguzo nyingi za ukubwa unaotaka kadri kontena inavyoruhusu. Unaweza kutengeneza mipangilio inayobadilika unapochanganya `auto-fill` na `minmax`, kama ifuatavyo:

```css
repeat(auto-fill, minmax(60px, 1fr));
```

Wakati kontena inabadilika ukubwa, mpangilio huu unaendelea kuingiza nguzo za 60px na kuzieneza hadi iweze kuingiza nyingine zaidi. **Kumbuka:** Ikiwa kontena yako haiwezi kuingiza vitu vyote kwenye safu moja, itavihamisha chini kwenye safu mpya.

# --instructions--

Katika gridi ya kwanza, tumia `auto-fill` na `repeat` kujaza gridi na nguzo zenye upana wa kiwango cha chini `60px` na kiwango cha juu `1fr`. Kisha badilisha ukubwa wa onyesho kuona auto-fill ikifanya kazi.

# --hints--

Darasa la `container` linapaswa kuwa na sifa ya `grid-template-columns` yenye `repeat` na `auto-fill` zitakazojaza gridi na nguzo zenye upana wa kiwango cha chini `60px` na kiwango cha juu `1fr`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?auto-fill\s*?,\s*?minmax\s*?\(\s*?60px\s*?,\s*?1fr\s*?\)\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(3, minmax(60px, 1fr));

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }

  .container2 {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: Silver;
    display: grid;
    grid-template-columns: repeat(3, minmax(60px, 1fr));
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>
<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
<div class="container2">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(auto-fill, minmax(60px, 1fr));

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }

  .container2 {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: Silver;
    display: grid;
    grid-template-columns: repeat(3, minmax(60px, 1fr));
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>
<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
<div class="container2">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```
