---
id: 5a94fe2669fb03452672e45e
title: Tumia grid-area bila kuunda kiolezo cha maeneo
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c6N7VhK'
forumTopicId: 301135
dashedName: use-grid-area-without-creating-an-areas-template
---

# --description--

Sifa ya `grid-area` uliyojifunza katika changamoto iliyopita inaweza kutumika kwa njia nyingine. Ikiwa gridi yako haina kiolezo cha maeneo cha kurejelea, unaweza kuunda eneo kwa haraka kwa kipengee kuwekwa kama ifuatavyo:

```css
item1 {
  grid-area: 1/1/2/4;
}
```

Hii inatumia nambari za mistari uliyojifunza awali kuainisha wapi eneo la kipengee hiki litakuwa. Nambari katika mfano hapo juu zinawakilisha thamani hizi:

```css
grid-area: horizontal line to start at / vertical line to start at / horizontal
  line to end at / vertical line to end at;
```

Hivyo kipengee katika mfano kitachukua safu kati ya mistari ya 1 na 2, na nguzo kati ya mistari ya 1 na 4.

# --instructions--

Ukizingatia sifa ya `grid-area`, weka kipengee chenye darasa la `item5` kati ya mistari ya tatu na ya nne ya usawa na kati ya mistari ya kwanza na ya nne ya wima.

# --hints--

Darasa la `item5` linapaswa kuwa na sifa ya `grid-area` ili kujaza eneo lote kati ya mistari ya tatu na ya nne ya usawa, na mistari ya kwanza na ya nne ya wima.

```js
assert.match(
  code,
  /.item5\s*?{[\s\S]*grid-area\s*?:\s*?3\s*?\/\s*?1\s*?\/\s*?4\s*?\/\s*?4\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }

  .item5 {
    background: PaleGreen;
    /* Only change code below this line */


    /* Only change code above this line */
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item5 {
    grid-area: 3/1/4/4;
  }
</style>
```
