---
id: 5a94fe0569fb03452672e45c
title: Gawanya gridi kuwa eneo la template
challengeType: 0
forumTopicId: 301130
dashedName: divide-the-grid-into-an-area-template
---

# --description--

Unaweza kuunganisha seli za gridi yako pamoja kuwa <dfn>eneo</dfn> na kumpa eneo hilo jina la kipekee. Fanya hivi kwa kutumia `grid-template-areas` kwenye kontena kama ifuatavyo:

```css
grid-template-areas:
  'header header header'
  'advert content content'
  'advert footer footer';
```

Msimbo hapo juu unaunganisha seli za gridi kuwa maeneo manne: `header`, `advert`, `content`, na `footer`. Kila neno linawakilisha seli moja na kila jozi ya alama za nukuu zinawakilisha safu moja.

# --instructions--

Badilisha template ili eneo la `footer` lifunikie safu yote ya chini. Kueleza maeneo hayawezi kuleta athari yoyote ya kuona kwa sasa. Baadaye, utatumia kipengee kutumia eneo kuona jinsi inavyofanya kazi.

# --hints--

Darasa la `container` linapaswa kuwa na sifa ya `grid-template-areas` kama mfano lakini na eneo la `footer` likifunika safu yote ya chini.

```js
assert.match(
  __helpers.removeCssComments(code),
  /.container\s*?{[\s\S]*grid-template-areas\s*?:\s*?["|']\s*?header\s*?header\s*?header\s*?["|']\s*?["|']\s*?advert\s*?content\s*?content\s*?["|']\s*?["|']\s*?footer\s*?footer\s*?footer\s*?["|']\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    grid-template-areas:
    /* Only change code below this line */
      'header header header'
      'advert content content'
      'advert footer footer';
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;

    grid-template-areas:
      'header header header'
      'advert content content'
      'footer footer footer';
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```
