---
id: 5a94fe6269fb03452672e462
title: Unda mipangilio inayobadilika kwa kutumia auto-fit
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c3dPph8'
forumTopicId: 301127
dashedName: create-flexible-layouts-using-auto-fit
---

# --description--

`auto-fit` hufanya kazi karibu sawa na `auto-fill`. Tofauti pekee ni kwamba wakati ukubwa wa kontena unazidi ukubwa wa vitu vyote pamoja, `auto-fill` huendelea kuingiza safu au nguzo tupu na kusukuma vitu vyako pembeni, wakati `auto-fit` hukunjua safu au nguzo hizo tupu na kunyoosha vitu vyako ili viendane na ukubwa wa kontena.

**Kumbuka:** Ikiwa kontena lako haliwezi kuweka vitu vyote kwenye safu moja, litaweka chini kwenye safu mpya.

# --instructions--

Katika gridi ya pili, tumia `auto-fit` pamoja na `repeat` kujaza gridi na nguzo zenye upana wa kiwango cha chini wa `60px` na kiwango cha juu cha `1fr`. Kisha badilisha ukubwa wa onyesho kuona tofauti.

# --hints--

Darasa la `container2` linapaswa kuwa na sifa ya `grid-template-columns` yenye `repeat` na `auto-fit` itakayojaza gridi na nguzo zenye upana wa kiwango cha chini wa `60px` na kiwango cha juu cha `1fr`.

```js
assert.match(
  code,
  /.container2\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?auto-fit\s*?,\s*?minmax\s*?\(\s*?60px\s*?,\s*?1fr\s*?\)\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: repeat(auto-fill, minmax(60px, 1fr));
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }

  .container2 {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: Silver;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(3, minmax(60px, 1fr));

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
<div class="container2">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(auto-fill, minmax(60px, 1fr));
  }
  .container2 {
    grid-template-columns: repeat(auto-fit, minmax(60px, 1fr));
  }
</style>
```
