---
id: 67116d7584d0b469b14579c0
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

在 `library` 数组内创建另一个具有以下属性和值的对象：

| Property | Value   |
| ----------- | ------- |
| `title` | `"Atomic Habits"`|
| `author` | `"James Clear"`|
| `about` | `"A practical book about discarding bad habits and building good ones"`|
| `pages` | `320` |

# --hints--


你的 `library` 数组的第二项应该是一个对象。

```js
assert.isObject(library[1]);
```

你的对象应有一个 `title` 属性。

```js
assert.property(library[1], "title")
```

你的 `title` 属性的值应为 `"Atomic Habits"`。

```js
assert.propertyVal(library[1], "title", "Atomic Habits");
```

你的对象应具有 `author` 属性。

```js
assert.property(library[1], "author")
```

你的 `author` 属性的值应为 `"James Clear"`。

```js
assert.propertyVal(library[1], "author", "James Clear");
```

你的对象应具有一个 `about` 属性。

```js
assert.property(library[1], "about")
```

你的 `about` 属性应具有值 `"A practical book about discarding bad habits and building good ones"`。

```js
assert.propertyVal(library[1], "about", "A practical book about discarding bad habits and building good ones");
```

你的对象应具有 `pages` 属性。

```js
assert.property(library[1], "pages")
```

你的 `pages` 属性的值应为数字。

```js
assert.isNumber(library[1]?.pages)
```

你的 `pages` 属性应设置为 `320`。

```js
assert.propertyVal(library[1], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
--fcc-editable-region--
  
--fcc-editable-region--
];
```
