---
id: 67116c175a77cb64c5c10d49
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在之前的课程中，你学习了如何使用像 `map`、`filter` 和 `reduce` 这样的高阶函数。

在本次研讨会中，你将构建一个库管理应用，这将为你提供练习使用不同高阶函数的机会。

要开始，创建一个名为 `library` 的变量并赋值为空数组。

# --hints--

你应该创建一个 `library` 数组。

```js
assert.isArray(library)
```

你的 `library` 数组应为空。

```js
assert.isEmpty(library)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
