---
id: 67116d7584d0b469b14579bf
title: 步骤2
challengeType: 1
dashedName: step-2
---

# --description--

在 `library` 数组中，创建一个具有以下属性和值的对象：

| Property | Value   |
| ----------- | ------- |
| `title` | `"Your Next Five Moves: Master the Art of Business Strategy"`|
| `author` | `"Patrick Bet-David and Greg Dinkin"`|
| `about` | `"A book on how to plan ahead"`|
| `pages` | `320` |

# --hints--

你应该在 `library` 数组中有一个对象。

```js
assert.isObject(library[0])
```

你的对象应有一个 `title` 属性。

```js
assert.property(library[0], "title")
```

你的 `title` 属性的值应为 `"Your Next Five Moves: Master the Art of Business Strategy"`。

```js
assert.propertyVal(library[0], "title", "Your Next Five Moves: Master the Art of Business Strategy");
```

你的对象应具有 `author` 属性。

```js
assert.property(library[0], "author")
```

你的 `author` 属性应具有值 `"Patrick Bet-David and Greg Dinkin"`。

```js
assert.propertyVal(library[0], "author", "Patrick Bet-David and Greg Dinkin");
```

你的对象应具有一个 `about` 属性。

```js
assert.property(library[0], "about")
```

你的 `about` 属性应具有值 `"A book on how to plan ahead"`。

```js
assert.propertyVal(library[0], "about", "A book on how to plan ahead");
```

你的对象应具有 `pages` 属性。

```js
assert.property(library[0], "pages")
```

你的 `pages` 属性的值应为数字。

```js
assert.isNumber(library[0]?.pages)
```

你的 `pages` 属性应设置为 `320`。

```js
assert.propertyVal(library[0], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
--fcc-editable-region--
  
--fcc-editable-region--
];
```
