---
id: 68225ea2f1aae301fa4df0f4
title: 步骤 9
challengeType: 1
dashedName: step-9
---

# --description--

当 `array` 的结果正在处理时，最终期望的结果应是一个列出所有带有标题和作者的书籍的 `string`。

这就是 `join()` 方法的用武之地。在之前的课程中，你已经学到 `join()` 方法用于返回一个新建的字串，该字串是所有数组元素连接成一个单一字串，并在每个元素之间使用指定的分隔符。

这里是复习内容：

```js
const developers = ["Naomi", "Tom", "Jessica"];

const teamList = developers.join("\n");
console.log(teamList); 
// Naomi
// Tom
// Jessica
```

此分支中的分隔符是 `\n`，它表现为换行字符。 

将 `join()` 方法与分隔符 `\n` 链接到 `map()`。现在你应该在控制台看到一个字串，而不是结果数组。 

# --hints--

你的 `getBookInformation` 函数应返回一个字串。

```js
assert.isString(getBookInformation(library));
```

你的 `getBookInformation` 函数应返回一个表现最终结果的单个字串。

```js
const testLibrary = [
  {
    title: "Title A",
    author: "Author A",
    about: "About A",
    pages: 320,
  },
  {
    title: "Title B",
    author: "Author B",
    about: "About B",
    pages: 320,
  },
  {
    title: "Title C",
    author: "Author C",
    about: "About C",
    pages: 304,
  },
];

const expected = `Title A by Author A
Title B by Author B
Title C by Author C`;

assert.strictEqual(getBookInformation(testLibrary), expected);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: 'A practical book about discarding bad habits and building good ones',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "A book that emphasizes the importance of identifying and understanding one's adversaries to succeed in the business world",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: 'A book focusing on how to build an audience-driven business',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: 'A book about effective ways to lead a coffee bean lifestyle',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: 'A book on how to develop creativity and  innovation skills',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: 'A book about financial literacy, financial independence, and building wealth. ',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: 'A book on how to bootstrap a business',
    pages: 500,
  },
];

console.log("Books in the Library:\n");

--fcc-editable-region--
function getBookInformation(catalog) {
  return catalog.map(book => `${book.title} by ${book.author}`)
}
--fcc-editable-region--

console.log(getBookInformation(library));
```
