---
id: 68c2fb153760d075321c5da3
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Camperbot 已經創建了一個程式設計旅程的部落格頁面，但看起來該頁面存在一些存取問題和不良做法。你在這個工作坊中的任務是為 Camperbot 修正這些問題。

第一組錯誤與標題的使用有關。在之前的課程中，你學到最佳做法是每頁只使用一個 `h1`。但看起來 Camperbot 使用了好幾個 `h1`。 

保留 `<h1>Welcome to Camperbot's Blog</h1>` 不變，並將其他 `h1` 元素改為 `h2` 元素。

# --hints--

你不應該更改 `<h1>Welcome to Camperbot's Blog</h1>`。

```js
const heading = document.querySelectorAll("h1")[0];
assert.equal(heading?.textContent, "Welcome to Camperbot's Blog");
```

你應該將文字為 `My Journey Learning to Code` 的 `h1` 元素改為 `h2` 元素。

```js
const firstPost = document.getElementById("post1");
assert.equal(firstPost?.tagName, "H2");
assert.equal(firstPost?.textContent, "My Journey Learning to Code");
```

你應該將文字為 `Accessibility Matters` 的 `h1` 元素改為 `h2` 元素。

```js
const secondPost = document.getElementById("post2");
assert.equal(secondPost?.tagName, "H2");
assert.equal(secondPost?.textContent, "Accessibility Matters");
```

你應該將文字為 `What's Next?` 的 `h1` 元素改為 `h2` 元素。

```js
const thirdPost = document.getElementById("post3");
assert.equal(thirdPost?.tagName, "H2");
assert.equal(thirdPost?.textContent, "What's Next?");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Camperbot's Blog</title>
</head>
<body>
  <h1>Welcome to Camperbot's Blog</h1>

  <h2>Navigation</h2>
  <ul>
    <li><a href="#post1">My Journey</a></li>
    <li><a href="#post2">Accessibility</a></li>
    <li><a href="#post3">Next Steps</a></li>
  </ul>
--fcc-editable-region--
  <div>
    <h1 id="post1">My Journey Learning to Code</h1>
    <p>I started learning to code a few months ago and it's been a wild ride!</p>
    
    <h4>Early Challenges</h4>
    <p>At first, syntax was really confusing.</p>
    
    <h4>Breakthroughs</h4>
    <p>Eventually things started to click.</p>
  </div>

  <div>
    <h1 id="post2">Accessibility Matters</h1>
    <p>Today I learned that not everyone uses the web the same way I do.</p>

    <h5>Screen Readers</h5>
    <p>These tools help visually impaired users browse websites.</p>
  </div>

  <div>
    <h1 id="post3">What's Next?</h1>
    <p>I'm excited to dive into JavaScript and build interactive features!</p>

    <h3>Coming soon: My first JavaScript project!</h3>
    <p>Stay tuned for some exciting interactive blog features.</p>
  </div>
--fcc-editable-region--
  <h3>Contact Me</h3>
  <p>Email me at camperbot@blog.io</p>
</body>
</html>
```
