---
id: 66c8ba975ee7230e29f6c4ac
title: Hatua ya 2
challengeType: 1
dashedName: step-2
---

# --description--

Wakati mtumizi anastahili mkopo, utataka kuonyesha ujumbe kwake kwenye konsoli.

Kwa ajili hiyo, utaunda kitendakazi ndani ambacho kitakuwa na ukaguzi fulani utakaorudisha mkopo ambao mwombaji anastahili.

Tengeneza kitendakazi tupu cha `getLoanMessage` chenye vigezo `annualIncome` na `creditScore`.

# --hints--

Unapaswa kuunda kitendakazi cha `getLoanMessage`

```js
assert.isFunction(getLoanMessage)
```

Kitendakazi chako cha `getLoanMessage` kinapaswa kuwa na vigezo `annualIncome` na `creditScore`.

```js
assert.match(getLoanMessage?.toString(), /\(\s*annualIncome,\s*creditScore\s*\)|\(\s*creditScore,\s*annualIncome\s*\)/);
```

`getLoanMessage` inapaswa kuwa kitendakazi tupu.

```js
assert.match(getLoanMessage?.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

--fcc-editable-region--

--fcc-editable-region--
```
