---
id: 66c8ba41a77db20a93f9d7a1
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utakagua masharti ya JavaScript kwa kujenga programu ya kukagua sifa za mkopo.

Programu itakagua kama mtumizi anastahili mkopo wa duplex, gari, na kondomu kulingana na mapato yao ya mwaka na alama ya mkopo. Alama ya mkopo ni nambari inayowakilisha jinsi ulivyo mzuri katika kusimamia pesa ulizokopa.

Ili kuanza, tengeneza vigezo na thamani zifuatazo.

| Jina la Kigezo | Thamani   |
| ----------- | ------- |
| `minIncomeForDuplex`       | `60000` |
| `minCreditScoreForDuplex`   | `700`   |
| `minIncomeForCondo`   | `45000`   |
| `minCreditScoreForCondo`    | `680`   |
| `minIncomeForCar`   | `30000`   |
| `minCreditScoreForCar`   | `650`   |

# --hints--

Unapaswa kutengeneza kigezo cha `minIncomeForDuplex` kilichowekwa kuwa `60000`.

```js
assert.equal(minIncomeForDuplex, 60000);
```

Unapaswa kutengeneza `minCreditScoreForDuplex` kilichowekwa kuwa `700`.

```js
assert.equal(minCreditScoreForDuplex, 700);
```

Unapaswa kutengeneza kigezo cha `minIncomeForCar` kilichowekwa kuwa `30000`.

```js
assert.equal(minIncomeForCar, 30000);
```

Unapaswa kutengeneza kigezo cha `minCreditScoreForCar` kilichowekwa kuwa `650`.

```js
assert.equal(minCreditScoreForCar, 650);
```

Unapaswa kutengeneza kigezo cha `minIncomeForCondo` kilichowekwa kuwa `45000`.

```js
assert.equal(minIncomeForCondo, 45000);
```

Unapaswa kutengeneza kigezo cha `minCreditScoreForCondo` kilichowekwa kuwa `680`.

```js
assert.equal(minCreditScoreForCondo, 680);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
