---
id: 66c8ba975ee7230e29f6c4b0
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Sasa, unapaswa kukagua kama mwombaji anastahili mkopo wa gari tu. Ikiwa anastahili, rudisha mfuatano wa herufi `"You qualify for a car loan."`.

# --hints--

Kitendakazi chako cha `getLoanMessage` kinapaswa kurudisha mfuatano wa herufi ikiwa mwombaji anastahili mkopo wa gari.

```js
assert.isString(getLoanMessage(30000, 650));
```

Kitendakazi chako cha `getLoanMessage` kinapaswa kurudisha mfuatano wa herufi `"You qualify for a car loan."`.

```js
assert.strictEqual(getLoanMessage(30000, 650), "You qualify for a car loan.");
```

Kitendakazi chako cha `getLoanMessage` kinapaswa kurudisha `undefined` ikiwa mapato ya mwaka na alama ya mkopo ya mwombaji hayakidhi mahitaji ya mkopo wa gari.

```js
assert.strictEqual(getLoanMessage(30000, 600), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
--fcc-editable-region--
  } 
  
--fcc-editable-region--
}
```
