---
id: 66c8ba975ee7230e29f6c4ad
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Ili kuangalia ni mkopo gani mtumizi anastahili kulipwa kulingana na `annualIncome` na `creditScore`, lazima utumie tamko la `if/else if` au kikokotoo cha masharti cha pande tatu ndani ya kitendakazi cha `getLoanMessage`. Kisha utarudisha ujumbe unaofaa katika sehemu ya kila hali.

Kuanzia na mkopo wa duplex, angalia kama `annualIncome` ni kubwa au sawa na `minIncomeForDuplex` NA kama `creditScore` ni kubwa au sawa na `minCreditScoreForDuplex`.

Kama hali hiyo ni kweli, basi mwombaji anastahili mkopo wa duplex na mikopo mingine. Hivyo, ndani ya ukaguzi, rudisha mfuatano wa herufi `"You qualify for a duplex, condo, and car loan."`

# --hints--

Kitendakazi chako cha `getLoanMessage` kinapaswa kurudisha mfuatano wa herufi.

```js
assert.isString(getLoanMessage(65000, 750));
assert.isString(getLoanMessage(60000, 750));
assert.isString(getLoanMessage(65000, 700));
assert.isString(getLoanMessage(60000, 700));
```

Kitendakazi chako cha `getLoanMessage` kinapaswa kurudisha mfuatano wa herufi `"You qualify for a duplex, condo, and car loan."`.

```js
assert.strictEqual(getLoanMessage(65000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(65000, 700), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 700), "You qualify for a duplex, condo, and car loan.");
```

Kitendakazi chako cha `getLoanMessage` kinapaswa kurudisha `undefined` ikiwa mapato ya mwaka na alama ya mkopo ya mwombaji hayakidhi mahitaji ya mkopo wa duplex.

```js
assert.isUndefined(getLoanMessage(59000, 690));
assert.isUndefined(getLoanMessage(59000, 700));
assert.isUndefined(getLoanMessage(59000, 750));
assert.isUndefined(getLoanMessage(60000, 690));
assert.isUndefined(getLoanMessage(65000, 690));
```
 
# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
