---
id: 66c8ba975ee7230e29f6c4b1
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Ikiwa kipato cha mwaka cha mwombaji na alama ya mkopo ni chini ya `minIncomeForCar` na `minCreditScoreForCar`, basi hawastahili kwa mkopo wowote. Kwa hivyo, rudisha mfuatano wa herufi `"You don't qualify for any loans."`

# --hints--

Kitendakazi chako cha `getLoanMessage` kinapaswa kurudisha mfuatano wa herufi ikiwa mwombaji hastahili kwa mkopo wowote.

```js
assert.isString(getLoanMessage(25000, 550));
```

Kitendakazi chako cha `getLoanMessage` kinapaswa kurudisha mfuatano wa herufi `"You don't qualify for any loans."` ikiwa mwombaji hatimii vigezo vya mkopo wowote.

```js
assert.strictEqual(getLoanMessage(15000, 600), "You don't qualify for any loans.");
assert.strictEqual(getLoanMessage(25000, 550), "You don't qualify for any loans.");
assert.strictEqual(getLoanMessage(20000, 500), "You don't qualify for any loans.");
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
