---
id: 66c8ba975ee7230e29f6c4b2
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Sasa, ni wakati wa kujaribu kitendakazi chako cha `getLoanMessage`.

Tumia jedwali lililopo hapa chini kuunda vigezo vya `4` na thamani zao:

| Jina la Kigezo | Thamani   |
| ----------- | ------- |
| `duplexLoanMsg`       | `getLoanMessage(85000, 850)` |
| `condoLoanMsg`   | `getLoanMessage(65000, 690)`   |
| `carLoanMsg`   | `getLoanMessage(45000, 660)`   |
| `noLoanMsg`   | `getLoanMessage(25000, 550)`   |

Baada ya hapo, andika kila kigezo kwenye konsoli ili kuona ujumbe.

Kwa hivyo, mradi wako wa ukaguzi wa sifa za mkopo umekamilika!

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `duplexLoanMsg`.

```js
assert.isNotNull(duplexLoanMsg);
```

Kigezo chako cha `duplexLoanMsg` kinapaswa kuwekwa kwa matokeo ya `getLoanMessage(85000, 850)`.

```js
assert.equal(duplexLoanMsg, getLoanMessage(85000, 850));
```

Unapaswa kuandika kigezo cha `duplexLoanMsg` kwenye konsoli.

```js
assert.match(code, /console\.log\(duplexLoanMsg\)/)
```

Unapaswa kuwa na kigezo kinachoitwa `condoLoanMsg`.

```js
assert.isNotNull(condoLoanMsg);
```

Kigezo chako cha `condoLoanMsg` kinapaswa kuwekwa kwa matokeo ya `getLoanMessage(65000, 690)`.

```js
assert.equal(condoLoanMsg, getLoanMessage(65000, 690));
```

Unapaswa kuandika kigezo cha `condoLoanMsg` kwenye konsoli.

```js
assert.match(code, /console\.log\(condoLoanMsg\)/)
```

Unapaswa kuwa na kigezo kinachoitwa `carLoanMsg`.

```js
assert.isNotNull(carLoanMsg);
```

Kigezo chako cha `carLoanMsg` kinapaswa kuwekwa kwa matokeo ya `getLoanMessage(45000, 660)`.

```js
assert.equal(carLoanMsg, getLoanMessage(45000, 660));
```

Unapaswa kuandika kigezo cha `carLoanMsg` kwenye konsoli.

```js
assert.match(code, /console\.log\(carLoanMsg\)/)
```

Unapaswa kuwa na kigezo kinachoitwa `noLoanMsg`.

```js
assert.isNotNull(noLoanMsg);
```

Kigezo chako cha `noLoanMsg` kinapaswa kuwekwa kwa matokeo ya `getLoanMessage(25000, 550)`.

```js
assert.equal(noLoanMsg, getLoanMessage(25000, 550));
```

Unapaswa kuandika kigezo cha `noLoanMsg` kwenye konsoli.

```js
assert.match(code, /console\.log\(noLoanMsg\)/)
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

const duplexLoanMsg = getLoanMessage(85000, 850)
console.log(duplexLoanMsg)

const condoLoanMsg = getLoanMessage(65000, 690)
console.log(condoLoanMsg)

const carLoanMsg = getLoanMessage(45000, 660)
console.log(carLoanMsg)

const noLoanMsg = getLoanMessage(25000, 550)
console.log(noLoanMsg)
```
