---
id: 66c8ba975ee7230e29f6c4af
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Ikiwa kipato cha mwaka cha mwombaji ni kikubwa au sawa na `minIncomeForCondo`, NA ikiwa alama yao ya mkopo ni kubwa au sawa na `minCreditScoreForCondo`, basi wanastahili mkopo wa kondomu na gari.

Kagua kama hiyo ni kweli katika kitendakazi cha `getLoanMessage`. Ikiwa ni kweli, rudisha mfuatano wa herufi `"You qualify for a condo and car loan."`

# --hints--

Kitendakazi chako cha `getLoanMessage` kinapaswa kurudisha mfuatano wa herufi ikiwa mwombaji anastahili mkopo wa kondomu na gari.

```js
assert.isString(getLoanMessage(45000, 680));
```

Kitendakazi chako cha `getLoanMessage` kinapaswa kurudisha mfuatano wa herufi `"You qualify for a condo and car loan."`.

```js
assert.strictEqual(getLoanMessage(45000, 680), "You qualify for a condo and car loan.");
```

Kitendakazi chako cha `getLoanMessage` kinapaswa kurudisha `undefined` ikiwa kipato cha mwaka cha mwombaji na alama ya mkopo havikutimiza mahitaji ya mkopo wa kondomu.

```js
assert.strictEqual(getLoanMessage(45000, 650), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
