---
id: 587d7fb7367417b2b2512c0d
title: Tumia model.findById() kutafuta hifadhidata yako kwa _id
challengeType: 2
forumTopicId: 301544
dashedName: use-model-findbyid-to-search-your-database-by-id
---

# --description--

Wakati wa kuhifadhi hati, MongoDB huongeza moja kwa moja uwanja `_id`, na kuuweka kuwa ufunguo wa kipekee wa herufi na nambari. Kutafuta kwa `_id` ni operesheni inayojirudia mara nyingi sana, kwa hivyo Mongoose hutoa njia maalum kwa ajili yake.

# --instructions--

Badilisha `findPersonById` kutafuta mtu pekee aliye na `_id` fulani, ukitumia `Model.findById() -> Person`. Tumia hoja ya kitendakazi `personId` kama ufunguo wa kutafuta.

# --hints--

Kutafuta kipengee kwa Id kunapaswa kufanikiwa

```js
  const response = await fetch(code + '/_api/find-by-id');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'test', 'item.name is not what expected');
  assert.equal(data.age, 0, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['none'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```
