---
id: 587d7fb7367417b2b2512c0b
title: Tumia model.find() kutafuta katika hifadhidata yako
challengeType: 2
forumTopicId: 301543
dashedName: use-model-find-to-search-your-database
---

# --description--

Katika matumizi yake rahisi kabisa, `Model.find()` hupokea hati ya hoja ya utafutaji (kitu cha JSON) kama hoja ya kwanza, kisha wito wa kurudisha. Hurejesha safu ya data ya matokeo yanayolingana. Inasaidia aina nyingi sana za chaguzi za utafutaji. Soma zaidi katika nyaraka.

# --instructions--

Badilisha kitendakazi cha `findPeopleByName` ili upate watu wote wenye jina fulani, ukitumia <code>Model.find() -\> [Person]</code>

Tumia hoja ya kitendakazi `personName` kama ufunguo wa utafutaji.

# --hints--

Kupata vitu vyote vinavyolingana na vigezo inapaswa kufanikiwa

```js
  const response = await fetch(code + '/_api/find-all-by-name', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'r@nd0mN4m3',
      age: 24,
      favoriteFoods: ['pizza']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an Array');
  assert.equal(
    data[0].name,
    'r@nd0mN4m3',
    'item.name is not what expected'
  );
  assert.equal(data[0].__v, 0, 'The item should be not previously edited');
```
