---
id: 587d7fb6367417b2b2512c09
title: Unda na Hifadhi Rekodi ya Mfano
challengeType: 2
forumTopicId: 301536
dashedName: create-and-save-a-record-of-a-model
---

# --description--

Katika changamoto hii utalazimika kuunda na kuhifadhi rekodi ya mfano.

# --instructions--

Ndani ya kitendakazi cha `createAndSavePerson`, unda mfano wa hati ukitumia kijenzi cha mfano wa `Person` ulichounda awali. Pitia kijenzi hicho kitu chenye sehemu za `name`, `age`, na `favoriteFoods`. Aina zao lazima zilingane na zile zilizopo katika `personSchema`. Kisha, itisha njia ya `document.save()` kwenye mfano wa hati uliorejeshwa. Pitia kwa njia hiyo wito wa kurudisha ukitumia mkataba wa Nodi. Hii ni mtindo wa kawaida; njia zote za CRUD zinazofuata huchukua kitendakazi cha mwitiko kama hoja ya mwisho.

```js
/* Example */

// ...
person.save(function(err, data) {
  //   ...do your stuff here...
});
```

# --hints--

Kuunda na kuhifadhi kipengee cha hifadhidata kunapaswa kufanikiwa

```js
  const response = await fetch(code + '/_api/create-and-save-person');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isString(data.name, '"item.name" should be a String');
  assert.isNumber(data.age, '28', '"item.age" should be a Number');
  assert.isArray(
    data.favoriteFoods,
    '"item.favoriteFoods" should be an Array'
  );
  assert.equal(data.__v, 0, 'The db item should be not previously edited');
```
