---
id: 587d7fb8367417b2b2512c10
title: Futa hati moja kwa kutumia model.findByIdAndRemove
challengeType: 2
forumTopicId: 301539
dashedName: delete-one-document-using-model-findbyidandremove
---

# --description--

`findByIdAndRemove` na `findOneAndRemove` ni kama njia za sasisho zilizotangulia. Zinapita hati iliyofutwa kwa hifadhidata. Kama kawaida, tumia hoja ya kitendakazi `personId` kama ufunguo wa utafutaji.

# --instructions--

Badilisha kitendakazi cha `removeById` kufuta mtu mmoja kwa `_id` ya mtu huyo. Unapaswa kutumia mojawapo ya njia `findByIdAndRemove()` au `findOneAndRemove()`.

# --hints--

Kufuta kipengee kunapaswa kufanikiwa

```js
  const response = await fetch(code + '/_api/remove-one-person', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Jason Bourne',
      age: 36,
      favoriteFoods: ['apples']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Jason Bourne', 'item.name is not what expected');
  assert.equal(data.age, 36, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['apples'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0);
  assert.equal(data.count, 0, 'the db items count is not what expected');
```
