---
id: 6619240f46cec8e04d77e03a
title: 基本的な関数の練習A
challengeType: 1
dashedName: top-basic-functions-exercise-a
---

# --description--

整数を受け取る関数を作成してください。この関数は、整数が`integer + 7`より小さい場合に与えられた`10`を返す必要があります。整数が`10`以上の場合は、与えられた`integer - 3`を返す必要があります。

関数の名前は`addOrSubtract`にしてください。

# --hints--

`addOrSubtract`という名前の関数を作成してください。

```js
assert.isFunction(addOrSubtract);
```

関数は整数を引数として受け取る必要があります。

```js
assert.match(addOrSubtract.toString(), /\s*addOrSubtract\(\s*\w+\s*\)/);
```

整数が10より小さい場合は、与えられた整数に7を足した値を返してください。

```js
assert.strictEqual(addOrSubtract(5), 12);
```

整数が10以上の場合は、与えられた整数から3を引いた値を返してください。

```js
assert.strictEqual(addOrSubtract(10), 7);
```




# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addOrSubtract(num) {
  if (num < 10) {
    return num + 7;
  } else {
    return num - 3;
  }
}
```
