---
id: 661e151f068359c3ccf2f4d7
title: 基本的な関数の練習 C
challengeType: 1
dashedName: top-basic-functions-exercise-c
---

# --description--

文字列をパラメータとして受け取り、最初の文字を大文字にした新しい文字列を返す、`capitalize` という名前の関数を書いてください。

# --hints--

`capitalize` という名前の関数を用意してください。

```js
assert.isFunction(capitalize);
```

関数は文字列をパラメータとして受け取る必要があります。

```js
assert.match(capitalize.toString(), /\s*capitalize\(\s*\w+\s*\)/);
```

関数は最初の文字を大文字にした新しい文字列を返す必要があります。

```js
assert.strictEqual(capitalize('sem'), 'Sem');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function capitalize(str) {
  return str[0].toUpperCase() + str.slice(1);
}
```
