---
id: 661e131f068359c3ccf2f4d6
title: 基本的な関数の練習 B
challengeType: 1
dashedName: top-basic-functions-exercise-b
---

# --description--

`multiply` という名前の関数を作成し、2つのパラメータを受け取り、それらの積を返してください。

# --hints--

`multiply` という名前の関数があるはずです。

```js
assert.isFunction(multiply);
```

関数は2つの整数を引数として受け取る必要があります。

```js
assert.match(multiply.toString(), /\s*multiply\(\s*\w+\s*,\s*\w+\s*\)/);
```

2つの整数の積を返す必要があります。

```js
assert.strictEqual(multiply(10, 10), 100);
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js 
function multiply(a, b) {
  return a * b;
}
```
