---
id: 661e17c6068359c3ccf2f4d8
title: 基本的な関数の練習 D
challengeType: 1
dashedName: top-basic-functions-exercise-d
---

# --description--

`lastLetter` という名前の関数を作成してください。この関数は文字列をパラメータとして受け取り、その文字列の最後の文字を返します。

# --hints--

`lastLetter` という名前の関数を作成してください。

```js
assert.isFunction(lastLetter);
```

関数は文字列をパラメータとして受け取る必要があります。

```js
assert.match(lastLetter.toString(), /\s*lastLetter\(\s*\w+\s*\)/);
```

文字列の最後の文字を返してください。

```js
assert.strictEqual(lastLetter('Sem'), 'm');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function lastLetter(str) {
  return str[str.length - 1];
}
```
