---
id: 694a9ff65d4472dd451d8255
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Neste workshop, você praticará como usar booleanos e declarações condicionais em Python construindo uma calculadora de reserva de ingressos de cinema.

Comece declarando variáveis que armazenam informações sobre o usuário e a exibição do filme.

Como você aprendeu nas lições anteriores, variáveis são atribuídas usando o operador de atribuição (`=`) desta forma:

```py
# Assigns the value 10 to variable x
x = 10
```

Crie uma variável chamada `base_price` para armazenar o preço base do ingresso do filme e defina seu valor como `15`. Crie outra variável chamada `age` para armazenar a idade do usuário e defina seu valor como `21`.

# --hints--

Você deve declarar uma variável chamada `base_price`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("base_price")`)) })
```

Você deve atribuir o valor inteiro `15` a `base_price`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("base_price").is_equivalent("base_price = 15")
`)) })
```

Você deve declarar uma variável chamada `age`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("age")`)) })
```

Você deve atribuir o valor inteiro `21` a `age`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("age").is_equivalent("age = 21")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
