---
id: 694cc1b515976889afd4cded
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Agora você precisa armazenar alguns valores de string sobre o ingresso do filme. Você pode armazenar strings em Python desta forma:

```py
name = 'Alice'
```

Crie uma variável chamada `seat_type` para armazenar o tipo de assento que o usuário selecionou e defina seu valor para a string `Gold`. Crie outra variável chamada `show_time` para armazenar o horário do filme e defina seu valor para a string `Evening`.

Lembre-se de cercar ambos os valores com aspas simples ou duplas.

# --hints--

Você deve declarar uma variável chamada `seat_type`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("seat_type")`)) })
```

Você deve atribuir a string `Gold` a `seat_type`. Lembre-se de cercar o valor com aspas simples ou duplas.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("seat_type").is_equivalent("seat_type = 'Gold'")
`)) })
```

Você deve declarar uma variável chamada `show_time`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("show_time")`)) })
```

Você deve atribuir a string `Evening` a `show_time`. Lembre-se de cercar o valor com aspas simples ou duplas.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("show_time").is_equivalent("show_time = 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
--fcc-editable-region--

--fcc-editable-region--
```
