---
id: 6958dfa0c74b237e1b56ce04
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

O usuário se qualifica para um desconto de associação se ele for um membro.

Crie uma variável chamada `discount` e defina seu valor como `0`. Isso armazenará o desconto que o usuário recebe no ingresso do filme.

# --hints--

Você deve declarar uma variável chamada `discount`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("discount")`)) })
```

Você deve atribuir o valor inteiro `0` a `discount`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("discount").is_equivalent("discount = 0")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

--fcc-editable-region--

--fcc-editable-region--
```
