---
id: 6958dd85587e10c99162dd81
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Como você aprendeu em lições anteriores, em Python, uma declaração `if` pode ser usada para executar código dependendo se uma condição é verdadeira.

Uma declaração `if` consiste na palavra-chave `if`, seguida por uma condição e dois pontos. O código a ser executado quando a condição for verdadeira, que deve ser indentado, é chamado de corpo da declaração `if`.

```py
if condition:
    # Code to execute if condition is True
```

Nesta etapa, você verificará se o usuário está elegível para reservar um ingresso de cinema com base na idade dele.

Crie uma declaração `if` para verificar se `age` é maior que `17`. Dentro do corpo da declaração `if`, imprima `User is eligible to book a ticket`. Isso imprimirá a mensagem somente quando a idade do usuário for maior que `17`.

Lembre-se de indentar o corpo da instrução `if` e cercar a mensagem com aspas simples ou duplas dentro da chamada `print()`.

# --hints--

Você deve ter uma instrução `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[0]`)) })
```

A condição da instrução `if` deve ser `age > 17`.

```js 
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("age > 17") or _cond.is_equivalent("17 < age")
`)) })
```

Você deve ter uma chamada `print()` dentro do seu `if`. Lembre-se de indentar o corpo do `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].find_calls("print")[0]
`)) })
```

Você deve imprimir `User is eligible to book a ticket` dentro do seu comando `if`. Lembre-se de cercar a mensagem com aspas simples ou duplas.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].has_call("print('User is eligible to book a ticket')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

--fcc-editable-region--

--fcc-editable-region--
```
