---
id: 6958e59845777391b86d7331
title: Passo 18
challengeType: 20
dashedName: step-18
---

# --description--

Em Python, uma declaração `if` também pode ser colocada dentro do corpo de outra declaração `if`. Isso é chamado de declaração `if` aninhada.

Uma declaração `if` aninhada permite que você verifique uma condição adicional somente depois que a primeira condição já foi satisfeita. A declaração `if` interna será executada somente se a condição `if` externa for verdadeira.

```py
if condition1:
    # Code to execute if condition1 is True
    if condition2:
        # Code to execute if both conditions are True
```

Agora você calculará as taxas de serviço com base no tipo de assento que o usuário selecionou.

Ainda dentro do corpo do seu último `if`, crie uma variável chamada `service_charges` e defina-a como `0`.

Então, crie uma declaração `if` aninhada para verificar se `seat_type` é igual a `Premium`. Dentro do corpo da declaração `if`, atualize o valor de `service_charges` para `5`.

# --hints--

Você deve declarar uma variável chamada `service_charges` dentro do corpo do seu último `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_variable("service_charges")
`)) })
```

Você deve atribuir o valor inteiro `0` a `service_charges`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 0")
`)) })
```

Você deve ter uma declaração `if` dentro do corpo da sua última declaração `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_ifs()[0]`)) })
```

A nova condição da declaração `if` deve ser `seat_type == 'Premium'`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("seat_type == 'Premium'")
`)) })
```

Você deve atribuir o valor inteiro `5` à variável `service_charges` dentro do seu novo comando `if`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_ifs()[0].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 5")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

if age >= 21 or age >= 18 and (show_time != 'Evening' or is_member):
    print('Ticket booking condition satisfied')

--fcc-editable-region--
    
--fcc-editable-region--
else:
    print('Ticket booking failed due to restrictions')
```
