---
id: 6958de0d773289a89515a6b8
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Agora você verificará se o usuário tem permissão para reservar um show noturno com base na idade dele.

Crie uma declaração `if` para verificar se `age` é maior ou igual a `21`. Dentro do corpo da declaração `if`, imprima `User is eligible for Evening shows`.

# --hints--

Você deve ter duas declarações `if` no seu código.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1]`)) })
```

A condição do seu segundo `if` deve ser `age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[1].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

Você deve imprimir `User is eligible for Evening shows` dentro do seu novo comando `if`. Lembre-se de cercar a mensagem com aspas simples ou duplas.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[0].has_call("print('User is eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

--fcc-editable-region--

--fcc-editable-region--
```
