---
id: 6958e1875e592b58e8852757
title: Passo 12
challengeType: 20
dashedName: step-12
---

# --description--

Agora crie uma variável chamada `extra_charges` e defina-a como `0`. Isso representará cobranças extras a serem aplicadas ao ingresso do filme nos finais de semana.

Então, crie uma declaração `if` para verificar se `is_weekend` é verdadeiro. Dentro do corpo da declaração `if`, atualize o valor de `extra_charges` para `2` e imprima `Extra charges will be applied` no terminal.

# --hints--

Você deve declarar uma variável chamada `extra_charges`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("extra_charges")`)) })
```

Você deve atribuir o valor inteiro `0` a `extra_charges`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("extra_charges").is_equivalent("extra_charges = 0")
`)) })
```

Você deve ter uma quarta declaração `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3]`)) })
```

A condição do seu quarto `if` deve ser `is_weekend`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_conditions()[0].is_equivalent("is_weekend")
`)) })
```

Você deve atribuir o valor inteiro `2` à variável `extra_charges` dentro do seu novo comando `if`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[3].find_bodies()[0].find_variable("extra_charges") 
assert _var.is_equivalent("extra_charges = 2")
`)) })
```

Você deve imprimir `Extra charges will be applied` dentro do seu novo comando `if`. Lembre-se de cercar a mensagem com aspas simples ou duplas.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[0].has_call("print('Extra charges will be applied')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

--fcc-editable-region--

--fcc-editable-region--
```
