---
id: 6958e2377cacbc46a7f7383a
title: Passo 13
challengeType: 20
dashedName: step-13
---

# --description--

Agora, adicione uma cláusula `else` à sua declaração `if` e imprima `No extra charges will be applied` dentro do corpo do `else`. Isso imprimirá a mensagem somente quando a condição de cobranças extras for falsa.

Então, abaixo da cláusula `else`, use a chamada `print()` para exibir uma mensagem que mostra `Extra charges:` seguida do valor atualizado de `extra_charges` e verifique a saída no terminal.

# --hints--

Você deve ter uma cláusula `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3].find_conditions()[1].is_empty()`)) })
```

Você deve imprimir `No extra charges will be applied` dentro da sua cláusula `else`. Lembre-se de cercar a mensagem com aspas simples ou duplas.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[1].has_call("print('No extra charges will be applied')")
`)) })
```

Você deve ter `print('Extra charges:', extra_charges)` abaixo da sua declaração `if...else`.

```js
({ test: () => (runPython(`
if_stmt = """if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')"""
print_call = "print('Extra charges:', extra_charges)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
--fcc-editable-region--

--fcc-editable-region--
```
