---
id: 6958e0ecc121b0ae3dd9c5e0
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

Você também deve lidar com o caso em que o usuário não se qualifica para um desconto de associação.

Adicione uma cláusula `else` à sua declaração `if` e imprima `User does not qualify for membership discount` dentro do corpo do `else`.

Você pode imprimir uma mensagem seguida por uma variável assim:

```py
print('Total:', total)
```

Você também quer exibir o valor atualizado de `discount`. Abaixo da declaração `if...else`, use a chamada `print()` para exibir uma mensagem que mostre `Discount:` seguido do valor atualizado de `discount`. Depois, verifique a saída no terminal.

# --hints--

Você deve ter uma cláusula `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2].find_conditions()[1].is_empty()`)) })
```

Você deve imprimir `User does not qualify for membership discount` dentro da sua cláusula `else`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[1].has_call("print('User does not qualify for membership discount')")
`)) })
```

Você deve ter `print('Discount:', discount)` abaixo da sua declaração `if...else`.

```js
({ test: () => (runPython(`
if_stmt = """if is_member:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')"""
print_call = "print('Discount:', discount)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
if is_member:
    discount = 3
    print('User qualifies for membership discount')
--fcc-editable-region--

--fcc-editable-region--
```
