---
id: 6960b8632b9544d1f57295f4
title: Passo 16
challengeType: 20
dashedName: step-16
---

# --description--

Usuários entre `18` e `21` podem reservar ingressos, mas somente quando o `show_time` não for `Evening`.

Use o operador `and` para construir uma expressão que verifica se `age` é maior ou igual a `18` e `show_time` não é `Evening`. Use o operador `or` para combiná-lo com a condição existente da última instrução `if`. Não crie novas variáveis. 

# --hints--

Sua última condição da declaração `if` deve ser `age >= 21 or age >= 18 and show_time != 'Evening'`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
_first = _cond.is_equivalent("age >= 21 or age >= 18 and show_time != 'Evening'")
_second = _cond.is_equivalent("age >= 21 or 18 <= age and show_time != 'Evening'")
assert _first or _second
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--
if age >= 21:
--fcc-editable-region--
    print('Ticket booking condition satisfied')
else:
    print('Ticket booking failed due to restrictions')
```
