---
id: 6960b6a6ea06e27dcc7eef37
title: Passo 15
challengeType: 20
dashedName: step-15
---

# --description--

Agora você verificará se o usuário satisfaz as condições para reservar um ingresso de cinema. Usuários com idade `21` ou acima podem sempre reservar ingressos sem quaisquer restrições.

Crie uma declaração `if` para verificar se `age` é maior ou igual a `21`. Dentro do corpo da declaração `if`, imprima `Ticket booking condition satisfied` no terminal.

Então, adicione uma cláusula `else` à sua declaração `if` e imprima `Ticket booking failed due to restrictions` dentro do corpo do `else`.

# --hints--

Você deve ter uma quinta declaração `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4]`)) })
```

A condição do seu quinto `if` deve ser `age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

Você deve imprimir `Ticket booking condition satisfied` dentro do seu novo comando `if`. Lembre-se de cercar a mensagem com aspas simples ou duplas.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_call("print('Ticket booking condition satisfied')")
`)) })
```

Sua nova declaração `if` deve ter uma cláusula `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_conditions()[1] == _Node()`)) })
```

Você deve imprimir `Ticket booking failed due to restrictions` dentro da sua cláusula `else`. Lembre-se de cercar a mensagem com aspas simples ou duplas.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[1].has_call("print('Ticket booking failed due to restrictions')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--

--fcc-editable-region--
```
