---
id: 6960a5165f7efe52a0c0aa39
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

Em Python, o operador `and` é usado para verificar se múltiplas condições são verdadeiras. Veja como você pode usá-lo para combinar duas condições em uma declaração `if`:

```py
if condition1 and condition2:
    # Code to execute if all conditions are True
```

O desconto de associação deve ser aplicado apenas aos membros se a `age` for maior ou igual a `21`.

Use o operador `and` para combinar a condição existente da sua declaração `if` com outra condição que verifica se `age` é maior ou igual a `21`.

# --hints--

A condição do seu terceiro `if` deve ser `is_member and age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[2].find_conditions()[0] 
assert _cond.is_equivalent("is_member and age >= 21") or _cond.is_equivalent("is_member and 21 <= age")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--
if is_member:
--fcc-editable-region--
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
