---
id: 6960a848aacf859f7bb15447
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Agora altere o valor da variável `is_member` para `False` pois o usuário não é um membro.

Depois disso, você verá que o valor de `discount` agora permanece `0`, porque ambas as condições devem ser satisfeitas para que o desconto seja aplicado.

# --hints--

Você deve atualizar o valor da variável `is_member` para `False`. Não coloque o valor entre aspas.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--
is_member = True
--fcc-editable-region--
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
