---
id: 696099abb6acc539b27055a1
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

No passo anterior, você verificou uma condição usando uma declaração `if`. A cláusula `else` é usada para tratar o caso quando a condição é falsa. Aqui está a sintaxe de uma declaração `if...else`:

```py
if condition:
    # Code to execute if condition is True
else:
    # Code to execute if condition is False
```

Agora, adicione uma cláusula `else` à sua declaração `if` e imprima `User is not eligible for Evening shows` dentro do corpo do `else`. Isso será impresso somente quando a condição na declaração `if` for falsa.

# --hints--

Você deve ter uma cláusula `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1].find_conditions()[1].is_empty()`)) })
```

Você deve imprimir `User is not eligible for Evening shows` dentro da sua cláusula `else`. Lembre-se de cercar a mensagem com aspas simples ou duplas.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[1].has_call("print('User is not eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
--fcc-editable-region--

--fcc-editable-region--
```
