---
id: 6960b308bf41cc39c9b4b438
title: Passo 14
challengeType: 20
dashedName: step-14
---

# --description--

Em Python, o operador `or` é usado para verificar se pelo menos uma de duas condições é verdadeira. Veja como você pode usá-lo em uma declaração `if`:

```py
if condition1 or condition2:
    # Code to execute if any condition is True
```

Cobranças extras também devem ser aplicadas se o show for à noite. Use o operador `or` para combinar a condição existente da sua declaração `if` com uma segunda condição que verifica se `show_time` é igual à string `Evening`.

# --hints--

A condição do seu último `if` deve ser `is_weekend or show_time == 'Evening'`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_conditions()[0].is_equivalent("is_weekend or show_time == 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
--fcc-editable-region--
if is_weekend:
--fcc-editable-region--
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)
```
