---
id: 6960ba6e6c95ed4ad896c6f2
title: Passo 17
challengeType: 20
dashedName: step-17
---

# --description--

Há mais uma exceção às regras de reserva. Usuários entre `18` e `21` podem reservar shows noturnos se forem membros.

Quando múltiplos operadores lógicos são usados em uma declaração `if`, condições unidas com `and` são avaliadas antes das condições unidas com `or`. Parênteses `()` são usados em Python para agrupar condições e controlar a ordem em que elas são avaliadas.

```py
if condition1 and (condition2 or condition3):
    # Code to execute if conditions evaluate to True
else:
    # Code to execute if conditions evaluate to False
```

Agora, adicione outra condição à sua declaração `if` existente usando o operador `or` para verificar se `is_member` é verdadeiro. Use os parênteses `()` para agrupar as condições `show_time != 'Evening'` e `is_member` juntas conforme mostrado no exemplo acima.

# --hints--

A condição do seu último `if` deve ser `age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--
if age >= 21 or age >= 18 and show_time != 'Evening':
--fcc-editable-region--
    print('Ticket booking condition satisfied')
else:
    print('Ticket booking failed due to restrictions')
```
