---
id: 696099c1ac72a4e585de1d5b
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Algumas informações só podem ser verdadeiras ou falsas. Como você aprendeu em lições anteriores, isso pode ser representado usando valores booleanos.

Crie uma variável chamada `is_member` para indicar se o usuário é um membro e defina seu valor como `True`.

Abaixo da variável `is_member` crie outra variável chamada `is_weekend` para indicar se a exibição do filme é no fim de semana e defina seu valor como `False`. Não coloque o valor entre aspas.

# --hints--

Você deve declarar uma variável chamada `is_member`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_member")`)) })
```

Você deve atribuir o valor booleano `True` a `is_member`. Não coloque o valor entre aspas.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = True")
`)) })
```

Você deve declarar uma variável chamada `is_weekend`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_weekend")`)) })
```

Você deve atribuir o valor booleano `False` a `is_weekend`. Não coloque o valor entre aspas.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_weekend").is_equivalent("is_weekend = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--

--fcc-editable-region--
```
