---
id: 696cb889fc4fdeb337e52864
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

Em Python, condições `if` não precisam comparar valores explicitamente com `True` ou `False`. Em vez disso, elas dependem da veracidade dos valores. Valores verdadeiros avaliam para `True` em um contexto booleano, como a condição de uma declaração `if`. Estes incluem números diferentes de zero, strings não vazias e o valor booleano `True`.

```py
if value:
   print('value is truthy')
```

Crie uma declaração `if` para verificar se `is_member` é verdadeiro. Dentro do corpo da declaração `if`, atualize o valor de `discount` para `3` e imprima `User qualifies for membership discount` no terminal.

# --hints--

Você deve ter três declarações `if` no seu código.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2]`)) })
```

A condição do seu terceiro `if` deve ser `is_member`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_conditions()[0].is_equivalent("is_member")
`)) })
```

Você deve atribuir o valor inteiro `3` à variável `discount` dentro do seu novo comando `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].find_variable("discount").is_equivalent("discount = 3")
`)) })
```

Você deve imprimir `User qualifies for membership discount` dentro do seu novo comando `if`. Lembre-se de cercar a mensagem com aspas simples ou duplas.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].has_call("print('User qualifies for membership discount')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--

--fcc-editable-region--
```
