---
id: a0b5010f579e69b815e7c5d6
title: 스마트한 단어 교체 함수 만들기
challengeType: 26
dashedName: build-a-smart-word-replacement-function
---

# --description--

이 실습에서는 주어진 인수를 사용해 문장에서 검색 및 교체를 수행한 후 새 문장을 반환하는 함수를 만드세요.

**참고:** 교체할 때 원래 단어의 첫 글자 대소문자를 유지하세요. 예를 들어 `Book` 단어를 `dog` 단어로 교체하려면 `Dog` 형태로 교체해야 합니다.

**목표**: 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `myReplace`라는 이름의 함수를 만드세요.
1. `myReplace` 함수는 문자열, 교체할 단어, 교체할 단어 세 가지 인수를 받아야 합니다.
1. `myReplace` 함수는 교체된 새 문자열을 반환해야 합니다.
1. 교체할 때 원래 단어의 첫 글자 대소문자를 유지해야 합니다.

# --hints--

`myReplace`라는 이름의 함수를 만드세요.

```js
assert.isFunction(myReplace);
```

`myReplace`는 세 개의 인수를 받아야 합니다.

```js
assert.lengthOf(myReplace, 3);
```

`myReplace("Let us go to the store", "store", "mall")`는 문자열 `Let us go to the mall`를 반환해야 합니다.

```js
assert.deepEqual(
  myReplace('Let us go to the store', 'store', 'mall'),
  'Let us go to the mall'
);
```

`myReplace("He is Sleeping on the couch", "Sleeping", "sitting")`는 문자열 `He is Sitting on the couch`를 반환해야 합니다.

```js
assert.deepEqual(
  myReplace('He is Sleeping on the couch', 'Sleeping', 'sitting'),
  'He is Sitting on the couch'
);
```

`myReplace("I think we should look up there", "up", "Down")`는 문자열 `I think we should look down there`를 반환해야 합니다.

```js
assert.deepEqual(
  myReplace('I think we should look up there', 'up', 'Down'),
  'I think we should look down there'
);
```

`myReplace("This has a spellngi error", "spellngi", "spelling")`는 문자열 `This has a spelling error`를 반환해야 합니다.

```js
assert.deepEqual(
  myReplace('This has a spellngi error', 'spellngi', 'spelling'),
  'This has a spelling error'
);
```

`myReplace("His name is Tom", "Tom", "john")`는 문자열 `His name is John`를 반환해야 합니다.

```js
assert.deepEqual(
  myReplace('His name is Tom', 'Tom', 'john'),
  'His name is John'
);
```

`myReplace("Let us get back to more Coding", "Coding", "algorithms")`는 문자열 `Let us get back to more Algorithms`를 반환해야 합니다.

```js
assert.deepEqual(
  myReplace('Let us get back to more Coding', 'Coding', 'algorithms'),
  'Let us get back to more Algorithms'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function myReplace(str, before, after) {
  if (before.charAt(0) === before.charAt(0).toUpperCase()) {
    after = after.charAt(0).toUpperCase() + after.substring(1);
  } else {
    after = after.charAt(0).toLowerCase() + after.substring(1);
  }
  return str.replace(before, after);
}
```
