---
id: 669e2f60e83c011754f711f9
title: 构建一个旅行社页面
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个 `DOCTYPE` 声明。
1. 你应该有一个 `html` 元素，并且 `lang` 属性设置为 `en`。
1. 你应该有一个 `head` 元素，里面包含一个 `meta` 空元素，`charset` 属性设置为 `utf-8`，还有一个 `title` 元素，文本为 `Travel Agency Page`。
1. 你应该在 `head` 元素中有一个 `meta` 标签，包含你网站的简短描述以便 SEO。
1. 你应该有一个 `h1` 元素来展示你的旅游目的地。它应该是页面上唯一的 `h1` 元素。
1. 你应该在 `h1` 元素下方有一个段落，介绍旅游机会。
1. 你应该有一个 `h2` 元素，文本为 `Packages`。
1. 你应该有一个 `p` 元素，简要介绍各种套餐。
1. 你应该有一个无序列表元素，包含两个列表项。这两个列表项的文本分别是 `Group Travels` 和 `Private Tours`。每个列表项的文本都应该被一个锚点元素包裹。
1. 你应该有一个 `h2` 元素，文本为 `Top Itineraries`。
1. 你应该至少有三个 `figure` 元素，每个都包含一个锚点元素和一个描述图像的 `figcaption` 元素。
1. 这三个锚点元素应该包含一个带有合适 `alt` 属性和有效 `src` 属性的 `img` 元素作为内容。你可以使用 `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`、`https://cdn.freecodecamp.org/curriculum/labs/alps.jpg` 和 `https://cdn.freecodecamp.org/curriculum/labs/sea.jpg`。
1. 你所有五个锚点元素都应该有一个 `href` 属性，值为 `https://www.freecodecamp.org/learn`，并且有一个 `target` 属性，值为 `_blank`。

# --hints--

你的旅行社页面该有一个 `!DOCTYPE html` 声明。

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

你应该有一个 `html` 设置为 `lang` 的 `en` 元素。

```js
assert.match(code, /<\s*html(?=[\s>])[^>]*\slang\s*=\s*('|")en\1[^>]*>/i);
```

你的 `html` 元素应该有一个结束标签。

```js
assert.match(code, /<\/html\s*>/i);
```

你的 `head` 元素中应该有一个 `html` 元素。

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>/i
);
```

你的 `head` 元素应该有一个闭合标签。

```js
assert.match(code, /<\/head\s*>/i);
```

你应该在 `head` 元素内有一个 `title` 元素。

```js
assert.isNotNull(document.querySelector('head title'));
```

你的 `title` 元素应该有一个闭合标签。

```js
assert.match(code, /<\/title\s*>/i);
```

你的 `meta` 元素中应该有一两个个 `head` 元素。

```js
assert.isAtLeast(document.querySelectorAll('head meta').length, 2);
```

一个 `meta` 元素应该有一个值为 `name` 的 `description` 属性和一个非空的 `content` 属性。

```js
assert(
  Array.from(document.querySelectorAll('head meta')).some(
    (meta) =>
      meta.getAttribute('name')?.toLowerCase() === 'description' &&
      meta.getAttribute('content')?.trim()
  )
);
```

一个 `meta` 元素的 `charset` 属性应该设置为 `utf-8`。

```js
const charsetMeta = document.querySelector('head meta[charset]');
assert.equal(charsetMeta?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

你的 `title` 元素应该包含你的旅行社名称。

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

你的 `html` 元素中应该有一个 `body` 元素。

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>(?:(?!<\/html\s*>)[\s\S])*<\s*body\s*>/i
);
```

你的 `body` 元素应该有一个闭合标签。

```js
assert.match(code, /<\/body\s*>/i);
```

你应该有一个 `h1` 元素来展示你的旅行目的地。

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

你应该只有一个 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你应该在 `h1` 元素下面有一个 `p` 元素。

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

你的第一个 `p` 元素应该介绍旅行机会。

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

你的第一个 `h2` 元素应该包含文本 `Packages`。

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText.trim(), 'Packages');
```

你应该在第一个 `h2` 元素下面有一个 `p` 元素。

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

你的第二个 `p` 元素应该简要介绍各种套餐。

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

你应该在第二个 `p` 元素下面有一个无序列表元素。

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

你的无序列表中应该有两个项。

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

你的两个列表项都应该包含一个锚点元素。

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

第一个列表项的锚点元素应该包裹文本 `Group Travels`。

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText.trim(), 'Group Travels');
```

第二个列表项的锚点元素应该包裹文本 `Private Tours`。

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText.trim(), 'Private Tours');
```

你应该在无序列表之后有一个 `h2` 元素。

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

你的第二个 `h2` 元素应该包含文本 `Top Itineraries`。

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText.trim(), 'Top Itineraries');
```

你应该至少有三个 `figure` 元素。

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

每个 `figure` 元素的第一个嵌套元素应该是一个锚点元素。

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

每个 `figure` 元素的第二个嵌套元素应该是一个 `figcaption` 元素。

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

每个 `figcaption` 应该包含一些文本。

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

每个嵌套在 `figure` 元素中的 `a` 元素应该包含一个图像。

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

每个 `img` 元素应该有一个有效的 `src` 属性。

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

每个 `img` 元素应该有一个适当的 `alt` 属性。

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

每个 `a` 元素都应该有一个 `href` 属性，值为 `https://www.freecodecamp.org/learn`。不要忘记列表项中的链接。

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

每个 `a` 元素都应该有一个 `target` 属性，值为 `_blank`。不要忘记列表项中的链接。

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```
