---
id: 686b9720ee1d032bd77a480a
title: Debuggen eines ISBN-Validators
challengeType: 27
dashedName: debug-an-isbn-validator
---

# --description--

Die ISBN (International Standard Book Number) ist eine eindeutige Kennung, die kommerziellen Büchern zugewiesen wird. Sie kann entweder 10 oder 13 Ziffern lang sein, und die letzte Ziffer ist eine Prüfziffer, die aus den anderen Ziffern berechnet wird.

Camperbot hat versucht, einen eigenen ISBN-Validator zu erstellen. Dabei sind jedoch einige Fehler passiert.

In diesem Labor beheben Sie den vorhandenen Code, damit er richtig funktioniert.

Erwartetes Verhalten:
Wenn der Benutzer das Programm ausführt, wird die Eingabeaufforderung `Enter ISBN and length: ` angezeigt. Der Benutzer kann den ISBN-Code, den er validieren möchte, im Format `ISBN,length` eingeben. Der ISBN-Code sollte keine Bindestriche enthalten, gefolgt von seiner Länge (`10` oder `13`), getrennt durch ein Komma.

Beispiel-Eingaben:
`1530051126,10` für ISBN-10-Codes.
`9781530051120,13` für ISBN-13-Codes.

Wie man die Prüfziffer findet:

- Sie müssen die detaillierte Berechnungslogik in diesem Labor nicht kennen. Die Funktionen `calculate_check_digit_10` und `calculate_check_digit_13` übernehmen die Berechnung und geben die erwartete Prüfziffer als String zurück.
- Die Prüfziffer für ISBN-10-Codes kann eine Zahl von `0` bis `9` oder ein Großbuchstabe `X` sein.
- Die Prüfziffer für ISBN-13-Codes kann eine Zahl von `0` bis `9` sein.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten den `IndentationError` im aktuellen Code beheben.
1. Selbst wenn der Benutzer keinen durch Komma getrennten Wert eingibt, sollte das Programm den `IndexError` ohne Absturz verarbeiten.
1. Wenn der Benutzer keinen durch Komma getrennten Wert eingibt, sollte er die Meldung `Enter comma-separated values.` in der Konsole sehen, und das Programm sollte beendet werden.
1. Selbst wenn der Benutzer einen nicht-numerischen Wert für die Länge eingibt, sollte das Programm den `ValueError` ohne Absturz verarbeiten.
1. Wenn der Benutzer einen nicht-numerischen Wert für die Länge eingibt, sollte er die Meldung `Length must be a number.` in der Konsole sehen, und das Programm sollte beendet werden.
1. Sie sollten den Off-by-one-Fehler in der `validate_isbn`-Funktion beheben.
1. Sie sollten den `TypeError` im aktuellen Code beheben, der auftritt, wenn der Benutzer einen gültigen ISBN-Code eingibt.
1. Sie sollten den `IndexError` im aktuellen Code beheben, wenn der Benutzer einen gültigen ISBN-Code eingibt.
1. Selbst wenn der Benutzer einen falschen ISBN-Code mit anderen als numerischen Zeichen eingibt, sollte das Programm den `ValueError` ohne Absturz verarbeiten.
1. Wenn der Benutzer einen falschen ISBN-Code mit anderen als numerischen Zeichen eingibt, sollte er die Meldung `Invalid character was found.` in der Konsole sehen.
1. Wenn der Benutzer `1530051126,10` eingibt, sollte er die Meldung `Valid ISBN Code.` sehen.
1. Wenn der Benutzer `9781530051120,13` eingibt, sollte er die Meldung `Valid ISBN Code.` sehen.

**Wichtig**: Sie müssen den Aufruf von `main()` im globalen Bereich auskommentieren, damit die Tests richtig ausgeführt werden.

Wenn Sie das Projekt abschließen, sollte der Benutzer je nach eingegebenen Werten die folgenden Meldungen sehen.

| ISBN-Code | Länge | Meldung | Beispiel-Eingabe |
|-----------|--------|---------|------------------|
| Gültig | Gültig | `Valid ISBN Code.` | `1530051126,10` |
| Ungültige Nummer | Gültig | `Invalid ISBN Code.` | `1530051125,10` |
| Entspricht nicht der angegebenen Länge oder leer gelassen | Gültig | `ISBN-10 code should be 10 digits long.` oder `ISBN-13 code should be 13 digits long.`, je nach eingegebener Länge. | `9781530051120,10` oder `1530051126,13` |
| Enthält nicht-numerische Zeichen (außer der Prüfziffer) | Gültig | `Invalid character was found.` | `15-0051126,10` |
| Beliebig | Ungültige Nummer | `Length should be 10 or 13.` | `1530051126,9` |
| Beliebig | Enthält nicht-numerische Zeichen oder leer gelassen | `Length must be a number.` | `1530051125,A` |
| Nicht durch Komma getrennt | Nicht durch Komma getrennt | `Enter comma-separated values.` | `1530051125` |

Sie können die folgenden Werte verwenden, um Ihren Code manuell zu testen, wenn Sie möchten.

Beispiel-Eingaben für gültige ISBN-10-Codes:
`1530051126,10`
`9971502100,10`
`080442957X,10`

Beispiel-Eingaben für gültige ISBN-13-Codes:
`9781530051120,13`
`9781947172104,13`

# --hints--

Sie sollten den Aufruf der Funktion `main` auskommentieren, damit die restlichen Tests richtig funktionieren.

```js
({
  test: () => runPython('assert not _Node(_code).has_call("main()")')
})
```

Sie sollten eine `validate_isbn`-Funktion haben.

```js
({
  test: () => runPython('assert _Node(_code).has_function("validate_isbn")')
})
```

Sie sollten eine `calculate_check_digit_10`-Funktion haben.

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_10")')
})
```

Sie sollten eine `calculate_check_digit_13`-Funktion haben.

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_13")')
})
```

Sie sollten eine `main`-Funktion haben.

```js
({
  test: () => runPython('assert _Node(_code).has_function("main")')
})
```

Wenn der Benutzer einen Wert eingibt, der kein durch Komma getrennter Wert ist, sollten Sie die Meldung `Enter comma-separated values.` in der Konsole sehen.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

Wenn der Benutzer einen nicht-numerischen Wert für die Länge eingibt, sollten Sie die Meldung `Length must be a number.` in der Konsole sehen.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,a"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

Wenn der Benutzer einen falschen ISBN-Code mit anderen als numerischen Zeichen eingibt, sollten Sie die Meldung `Invalid character was found.` in der Konsole sehen.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "978IS3OOSII2O,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

Wenn der Benutzer `1530051126,10` eingibt, sollten Sie die Meldung `Valid ISBN Code.` in der Konsole sehen.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Wenn der Benutzer `9781530051120,13` eingibt, sollten Sie die Meldung `Valid ISBN Code.` sehen.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Wenn der Benutzer `1530051125,10` eingibt, sollten Sie die Meldung `Invalid ISBN Code.` sehen.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid ISBN Code." in out

`)
}})
```

Wenn der Benutzer `9781530051120,10` eingibt, sollten Sie die Meldung `ISBN-10 code should be 10 digits long.` sehen.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-10 code should be 10 digits long." in out

`)
}})
```

Wenn der Benutzer `1530051126,13` eingibt, sollten Sie die Meldung `ISBN-13 code should be 13 digits long.` sehen.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-13 code should be 13 digits long." in out

`)
}})
```

Wenn der Benutzer `15-0051126,10` eingibt, sollten Sie die Meldung `Invalid character was found.` sehen.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "15-0051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

Wenn der Benutzer `1530051126,9` eingibt, sollten Sie die Meldung `Length should be 10 or 13.` sehen.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,9"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length should be 10 or 13." in out

`)
}})
```

Wenn der Benutzer `1530051125,A` eingibt, sollten Sie die Meldung `Length must be a number.` sehen.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,A"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

Wenn der Benutzer `1530051125` eingibt, sollten Sie die Meldung `Enter comma-separated values.` sehen.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

Wenn der Benutzer `9971502100,10` eingibt, sollten Sie die Meldung `Valid ISBN Code.` sehen.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9971502100,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Wenn der Benutzer `080442957X,10` eingibt, sollten Sie die Meldung `Valid ISBN Code.` sehen.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "080442957X,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Wenn der Benutzer `9781947172104,13` eingibt, sollten Sie die Meldung `Valid ISBN Code.` sehen.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781947172104,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

# --seed--

## --seed-contents--

```py
def validate_isbn(isbn, length):
    if len(isbn, length) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return
    main_digits = isbn[0:length]
    given_check_digit = isbn[length]
    main_digits_list = [int(digit) for digit in main_digits]
    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)
    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')
def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11
    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10
    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]
    length = int(values[1])
    if length == 10 or length == 13:
    validate_isbn(isbn, length)
    else:
    print('Length should be 10 or 13.')

main()
```

# --solutions--

```py
def validate_isbn(isbn, length):
    if len(isbn) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return

    check_digit_index = length - 1

    main_digits = isbn[0:check_digit_index]
    given_check_digit = isbn[check_digit_index]

    try:
        main_digits_list = [int(digit) for digit in main_digits]
    except ValueError as e:
        print('Invalid character was found.')
        return

    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)

    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')


def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11

    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10

    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]

    try:
        length = int(values[1])
    except ValueError as e:
        print('Length must be a number.')
        return
    except IndexError as e:
        print('Enter comma-separated values.')
        return

    if length == 10 or length == 13:
        validate_isbn(isbn, length)
    else:
        print('Length should be 10 or 13.')


# main()
```
