---
id: 6814d8e1516e86b171929de4
title: "チャレンジ 1: 母音のバランス"
challengeType: 28
dashedName: challenge-1
---

# --description--

文字列が与えられたとき、文字列の前半の母音の数が後半の母音の数と等しいかどうかを判定してください。

- 文字列には任意の文字が含まれる場合があります。
- `a`、`e`、`i`、`o`、および`u`の文字は、大文字・小文字を問わず母音とみなします。
- 文字列の文字数が奇数の場合は、中央の文字を無視してください。

# --hints--

`isBalanced("racecar")` は `true` を返すべきです。

```js
assert.isTrue(isBalanced("racecar"));
```

`isBalanced("Lorem Ipsum")` は `true` を返すべきです。

```js
assert.isTrue(isBalanced("Lorem Ipsum"));
```

`isBalanced("Kitty Ipsum")` は `false` を返すべきです。

```js
assert.isFalse(isBalanced("Kitty Ipsum"));
```

`isBalanced("string")` は `false` を返すべきです。

```js
assert.isFalse(isBalanced("string"));
```

`isBalanced(" ")` は `true` を返すべきです。

```js
assert.isTrue(isBalanced(" "));
```

`isBalanced("abcdefghijklmnopqrstuvwxyz")` は `false` を返すべきです。

```js
assert.isFalse(isBalanced("abcdefghijklmnopqrstuvwxyz"));
```

`isBalanced("123A#b!E&*456-o.U")` は `true` を返すべきです。

```js
assert.isTrue(isBalanced("123A#b!E&*456-o.U"));
```

# --seed--

## --seed-contents--

```js
function isBalanced(s) {

  return s;
}
```

# --solutions--

```js
function isBalanced(s) {
  const vowels = 'aeiou';
  const half = Math.floor(s.length / 2);

  let firstHalf = s.slice(0, half);
  let secondHalf = s.length % 2 === 0 ? s.slice(half) : s.slice(half + 1);

  const countVowels = str =>
    str
      .toLowerCase()
      .split('')
      .filter(c => vowels.includes(c))
      .length;

  return countVowels(firstHalf) === countVowels(secondHalf);
}
```
