---
id: 681cb1a2dab50c87ddb2e514
title: "チャレンジ 3: フィボナッチ数列"
challengeType: 28
dashedName: challenge-3
---

# --description--

フィボナッチ数列は、各数字が直前の2つの数字の合計になる数の列です。`0` と `1` から始めると、数列の最初の10個の数字は `0`、`1`、`1`、`2`、`3`、`5`、`8`、`13`、`21`、`34` です。

フィボナッチ数列の最初の2つの数字を含む配列と、数列の長さを表す整数が与えられたとき、その長さの数列を含む配列を返してください。

- 関数は長さが0以上の任意の長さの数列に対応する必要があります。
- 長さが0の場合は、空の配列を返してください。
- 最初の数字は数列の一部であることに注意してください。

# --hints--

`fibonacciSequence([0, 1], 20)` は `[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]` を返すべきです。

```js
assert.deepEqual(fibonacciSequence([0, 1], 20), [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]);
```

`fibonacciSequence([21, 32], 1)` は `[21]` を返すべきです。

```js
assert.deepEqual(fibonacciSequence([21, 32], 1), [21]);
```

`fibonacciSequence([0, 1], 0)` は `[]` を返すべきです。

```js
assert.deepEqual(fibonacciSequence([0, 1], 0), []);
```

`fibonacciSequence([10, 20], 2)` は `[10, 20]` を返すべきです。

```js
assert.deepEqual(fibonacciSequence([10, 20], 2), [10, 20]);
```

`fibonacciSequence([123456789, 987654321], 5)` は `[123456789, 987654321, 1111111110, 2098765431, 3209876541]` を返すべきです。

```js
assert.deepEqual(fibonacciSequence([123456789, 987654321], 5), [123456789, 987654321, 1111111110, 2098765431, 3209876541]);
```

# --seed--

## --seed-contents--

```js
function fibonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function fibonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2];
    sequence.push(nextValue);
  }
  return sequence;
}
```
