---
id: 681cb1afdab50c87ddb2e515
title: "チャレンジ 4: S  P  A  C  E  J  A  M"
challengeType: 28
dashedName: challenge-4
---

# --description--

文字列が与えられたら、その文字列からすべてのスペースを取り除き、各文字の間に2つのスペースを挿入し、すべてのアルファベットを大文字に変換して結果を返してください。

- アルファベット以外の文字は変更しないでください（スペースを除く）。

# --hints--

`spaceJam("freeCodeCamp")` は `"F  R  E  E  C  O  D  E  C  A  M  P"` を返すべきです。

```js
assert.equal(spaceJam("freeCodeCamp"), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("   free   Code   Camp   ")` は `"F  R  E  E  C  O  D  E  C  A  M  P"` を返すべきです。

```js
assert.equal(spaceJam("   free   Code   Camp   "), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("Hello World?!")` は `"H  E  L  L  O  W  O  R  L  D  ?  !"` を返すべきです。

```js
assert.equal(spaceJam("Hello World?!"), "H  E  L  L  O  W  O  R  L  D  ?  !");
```

`spaceJam("C@t$ & D0g$")` は `"C  @  T  $  &  D  0  G  $"` を返すべきです。

```js
assert.equal(spaceJam("C@t$ & D0g$"), "C  @  T  $  &  D  0  G  $");
```

`spaceJam("allyourbase")` は `"A  L  L  Y  O  U  R  B  A  S  E"` を返すべきです。

```js
assert.equal(spaceJam("all your base"), "A  L  L  Y  O  U  R  B  A  S  E");
```

# --seed--

## --seed-contents--

```js
function spaceJam(s) {

  return s;
}
```

# --solutions--

```js
function spaceJam(s) {
  return s.toUpperCase().replace(/\s+/g, '').split('').join('  ');
}
```
