---
id: 681cb1b0dab50c87ddb2e519
title: "チャレンジ 8: 階乗計算"
challengeType: 28
dashedName: challenge-8
---

# --description--

0から20までの整数が与えられたとき、その数の階乗を返してください。階乗とは、1からその数までのすべての数の積です。

- 0の階乗は1です。

# --hints--

`factorial(0)` は `1` を返す必要があります。

```js
assert.equal(factorial(0), 1);
```

`factorial(5)` は `120` を返す必要があります。

```js
assert.equal(factorial(5), 120);
```

`factorial(20)` は `2432902008176640000` を返す必要があります。

```js
assert.equal(factorial(20), 2432902008176640000);
```

# --seed--

## --seed-contents--

```js
function factorial(n) {

  return n;
}
```

# --solutions--

```js
function factorial(n) {
  return n == 0 ? 1 : n * factorial(n - 1);
}
```
