---
id: 681cb1afdab50c87ddb2e516
title: "チャレンジ 5: Jbelmud テキスト"
challengeType: 28
dashedName: challenge-5
---

# --description--

文字列が与えられたとき、各単語を次の条件で変換した文字列のシャッフル版を返してください。

- 単語の最初と最後の文字はそのままの位置に残す
- 最初と最後の文字の間のすべての文字はアルファベット順に並べ替える
- 入力される文字列には句読点が含まれず、すべて小文字です

# --hints--

`jbelmu("hello world")` は `"hello wlord"` を返すべきです。

```js
assert.equal(jbelmu("hello world"), "hello wlord");
```

`jbelmu("i love jumbled text")` は `"i love jbelmud text"` を返すべきです。

```js
assert.equal(jbelmu("i love jumbled text"), "i love jbelmud text");
```

`jbelmu("freecodecamp is my favorite place to learn to code")` は `"faccdeeemorp is my faiortve pacle to laern to cdoe"` を返すべきです。

```js
assert.equal(jbelmu("freecodecamp is my favorite place to learn to code"), "faccdeeemorp is my faiortve pacle to laern to cdoe");
```

`jbelmu("the quick brown fox jumps over the lazy dog")` は `"the qciuk borwn fox jmpus oevr the lazy dog"` を返すべきです。

```js
assert.equal(jbelmu("the quick brown fox jumps over the lazy dog"), "the qciuk borwn fox jmpus oevr the lazy dog");
```

# --seed--

## --seed-contents--

```js
function jbelmu(text) {

  return text;
}
```

# --solutions--

```js
function jbelmu(text) {
  return text
    .split(' ')
    .map((word) => {
      if (word.length <= 3) return word;
      const first = word[0];
      const last = word[word.length - 1];
      const middle = word
        .slice(1, -1)
        .split('')
        .sort()
        .join('');
      return first + middle + last;
    })
    .join(' ');
}
```
