---
id: 681cb1b0dab50c87ddb2e51a
title: "チャレンジ 9: 平方和"
challengeType: 28
dashedName: challenge-9
---

# --description--

1から指定された最大1,000までの正の整数に対して、そのすべての整数の平方の合計を返してください。

# --hints--

`sumOfSquares(5)` は `55` を返す必要があります。

```js
assert.equal(sumOfSquares(5), 55);
```

`sumOfSquares(10)` は `385` を返す必要があります。

```js
assert.equal(sumOfSquares(10), 385);
```

`sumOfSquares(25)` は `5525` を返す必要があります。

```js
assert.equal(sumOfSquares(25), 5525);
```

`sumOfSquares(500)` は `41791750` を返す必要があります。

```js
assert.equal(sumOfSquares(500), 41791750);
```

`sumOfSquares(1000)` は `333833500` を返す必要があります。

```js
assert.equal(sumOfSquares(1000), 333833500);
```

# --seed--

## --seed-contents--

```js
function sumOfSquares(n) {

  return n;
}
```

# --solutions--

```js
function sumOfSquares(n) {
  let total = 1;
  for(let i = 2; i <= n; i++) {
      total += i * i;
  }
  return total;
}
```
