---
id: 681cb1b0dab50c87ddb2e51b
title: "チャレンジ 10: 3ストライク"
challengeType: 28
dashedName: challenge-10
---

# --description--

1から10,000までの整数が与えられたとき、その整数までの数字のうち、平方の中に少なくとも1つの数字の3が含まれている数の個数を返してください。

# --hints--

`squaresWithThree(1)` は `0` を返すべきです。

```js
assert.equal(squaresWithThree(1), 0);
```

`squaresWithThree(10)` は `1` を返すべきです。

```js
assert.equal(squaresWithThree(10), 1);
```

`squaresWithThree(100)` は `19` を返すべきです。

```js
assert.equal(squaresWithThree(100), 19);
```

`squaresWithThree(1000)` は `326` を返すべきです。

```js
assert.equal(squaresWithThree(1000), 326);
```

`squaresWithThree(10000)` は `4531` を返すべきです。

```js
assert.equal(squaresWithThree(10000), 4531);
```

# --seed--

## --seed-contents--

```js
function squaresWithThree(n) {

  return n;
}
```

# --solutions--

```js
function squaresWithThree(n) {
  let count = 0;

  for (let i = 1; i <= n; i++) {
    const square = i * i;
    if (square.toString().includes('3')) {
      count++;
    }
  }

  return count;
}
```
