---
id: 681cb1b0dab50c87ddb2e518
title: "チャレンジ 7: 目標の合計"
challengeType: 28
dashedName: challenge-7
---

# --description--

数値の配列と整数の目標値が与えられたとき、その配列の中から目標値になるように合計できる2つの異なる数字を見つけてください。見つかった場合は、その2つの数字のインデックスを配列で返してください。もし合計が目標値になる2つの数字がなければ、`"Target not found"`を返してください。

- 返す配列のインデックスは昇順である必要があります。

# --hints--

`findTarget([2, 7, 11, 15], 9)`は`[0, 1]`を返すべきです。

```js
assert.deepEqual(findTarget([2, 7, 11, 15], 9), [0, 1]);
```

`findTarget([3, 2, 4, 5], 6)`は`[1, 2]`を返すべきです。

```js
assert.deepEqual(findTarget([3, 2, 4, 5], 6), [1, 2]);
```

`findTarget([1, 3, 5, 6, 7, 8], 15)`は`[4, 5]`を返すべきです。

```js
assert.deepEqual(findTarget([1, 3, 5, 6, 7, 8], 15), [4, 5]);
```

`findTarget([1, 3, 5, 7], 14)`は`"Target not found"`を返すべきです。

```js
assert.equal(findTarget([1, 3, 5, 7], 14), "Target not found");
```

# --seed--

## --seed-contents--

```js
function findTarget(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findTarget(arr, target) {
  for (let i = 0; i < arr.length; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] + arr[j] === target) {
        return [i, j];
      }
    }
  }
  return 'Target not found'
}
```
