---
id: 681cb05adab50c87ddb2e513
title: "チャレンジ 2: 基数チェック"
challengeType: 28
dashedName: challenge-2
---

# --description--

数値を表す文字列と、2から36までの整数の基数が与えられたとき、その数値がその基数で有効かどうかを判定してください。

- 文字列には数字と大文字または小文字の文字が含まれる場合があります。
- チェックは大文字小文字を区別しません。
- 与えられた基数で、すべての文字が有効な数字であれば、その数は有効です。
- さまざまな基数での有効な数字の例：
  - 基数 2: 0-1
  - 基数 8: 0-7
  - 基数 10: 0-9
  - 基数 16: 0-9 と A-F
  - 基数 36: 0-9 と A-Z

# --hints--

`isValidNumber("10101", 2)` は `true` を返すべきです。

```js
assert.isTrue(isValidNumber("10101", 2))
```

`isValidNumber("10201", 2)` は `false` を返すべきです。

```js
assert.isFalse(isValidNumber("10201", 2))
```

`isValidNumber("76543210", 8)` は `true` を返すべきです。

```js
assert.isTrue(isValidNumber("76543210", 8))
```

`isValidNumber("9876543210", 8)` は `false` を返すべきです。

```js
assert.isFalse(isValidNumber("9876543210", 8))
```

`isValidNumber("9876543210", 10)` は `true` を返すべきです。

```js
assert.isTrue(isValidNumber("9876543210", 10))
```

`isValidNumber("ABC", 10)` は `false` を返すべきです。

```js
assert.isFalse(isValidNumber("ABC", 10))
```

`isValidNumber("ABC", 16)` は `true` を返すべきです。

```js
assert.isTrue(isValidNumber("ABC", 16))
```

`isValidNumber("Z", 36)` は `true` を返すべきです。

```js
assert.isTrue(isValidNumber("Z", 36))
```

`isValidNumber("ABC", 20)` は `true` を返すべきです。

```js
assert.isTrue(isValidNumber("ABC", 20))
```

`isValidNumber("4B4BA9", 16)` は `true` を返すべきです。

```js
assert.isTrue(isValidNumber("4B4BA9", 16))
```

`isValidNumber("5G3F8F", 16)` は `false` を返すべきです。

```js
assert.isFalse(isValidNumber("5G3F8F", 16))
```

`isValidNumber("5G3F8F", 17)` は `true` を返すべきです。

```js
assert.isTrue(isValidNumber("5G3F8F", 17))
```

`isValidNumber("abc", 10)` は `false` を返すべきです。

```js
assert.isFalse(isValidNumber("abc", 10))
```

`isValidNumber("abc", 16)` は `true` を返すべきです。

```js
assert.isTrue(isValidNumber("abc", 16))
```

`isValidNumber("AbC", 16)` は `true` を返すべきです。

```js
assert.isTrue(isValidNumber("AbC", 16))
```

`isValidNumber("z", 36)` は `true` を返すべきです。

```js
assert.isTrue(isValidNumber("z", 36))
```

`isValidNumber("", 10)` は `false` を返すべきです。

```js
assert.isFalse(isValidNumber("", 10))
```

# --seed--

## --seed-contents--

```js
function isValidNumber(n, base) {

  return n;
}
```

# --solutions--

```js
function isValidNumber(n, base) {
  return new RegExp(`^[${'0123456789abcdefghijklmnopqrstuvwxyz'.substring(0, base)}]+\$`, "i").test(n);
}
```
