---
id: 681cb1afdab50c87ddb2e517
title: "チャレンジ 6: アナグラムチェッカー"
challengeType: 28
dashedName: challenge-6
---

# --description--

2つの文字列が与えられたとき、それらがアナグラム（同じ文字を任意の順序で含む）かどうかを判定してください。

- 大文字・小文字の区別と空白は無視してください。

# --hints--

`areAnagrams("listen", "silent")` は `true` を返すべきです。

```js
assert.isTrue(areAnagrams("listen", "silent"));
```

`areAnagrams("School master", "The classroom")` は `true` を返すべきです。

```js
assert.isTrue(areAnagrams("School master", "The classroom"));
```

`areAnagrams("A gentleman", "Elegant man")` は `true` を返すべきです。

```js
assert.isTrue(areAnagrams("A gentleman", "Elegant man"));
```

`areAnagrams("Hello", "World")` は `false` を返すべきです。

```js
assert.isFalse(areAnagrams("Hello", "World"));
```

`areAnagrams("apple", "banana")` は `false` を返すべきです。

```js
assert.isFalse(areAnagrams("apple", "banana"));
```

`areAnagrams("cat", "dog")` は `false` を返すべきです。

```js
assert.isFalse(areAnagrams("cat", "dog"));
```

# --seed--

## --seed-contents--

```js
function areAnagrams(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function areAnagrams(str1, str2) {
  const clean = (str) =>
    str.replace(/\s+/g, '').toLowerCase().split('').sort().join('');
    
  return clean(str1) === clean(str2);
}
```
