---
id: 6821ebea237de8297eaee795
title: "チャレンジ 19: キャンドルライト"
challengeType: 28
dashedName: challenge-19
---

# --description--

最初に持っているキャンドルの数を表す整数と、新しいキャンドルを作るのに必要な燃えたキャンドルの数を表す整数が与えられたとき、できるだけ多く作って燃やした後に使ったキャンドルの合計数を返してください。

例えば、7本のキャンドルがあり、新しいキャンドルを作るのに燃えたキャンドルが2本必要な場合：

1. 7本のキャンドルを燃やして7本の残りを得る、
2. 6本の残りをリサイクルして3本の新しいキャンドルを作る（1本の残りが残る）、
3. 3本のキャンドルを燃やしてさらに3本の残りを得る（合計4本）、
4. 4本の残りをリサイクルして2本の新しいキャンドルを作る、
5. 2本のキャンドルを燃やして2本の残りを得る、
6. 2本の残りをリサイクルして1本の新しいキャンドルを作る、
7. 1本のキャンドルを燃やす。

この例では合計で13本のキャンドルを燃やしたことになります。

# --hints--

`burnCandles(7, 2)` は `13` を返すべきです

```js
assert.equal(burnCandles(7, 2), 13);
```

`burnCandles(10, 5)` は `12` を返すべきです

```js
assert.equal(burnCandles(10, 5), 12);
```

`burnCandles(20, 3)` は `29` を返すべきです

```js
assert.equal(burnCandles(20, 3), 29);
```

`burnCandles(17, 4)` は `22` を返すべきです

```js
assert.equal(burnCandles(17, 4), 22);
```

`burnCandles(2345, 3)` は `3517` を返すべきです

```js
assert.equal(burnCandles(2345, 3), 3517);
```

# --seed--

## --seed-contents--

```js
function burnCandles(candles, leftoversNeeded) {

  return candles;
}
```

# --solutions--

```js
function burnCandles(candles, leftoversNeeded) {
  let totalBurned = 0;
  let unusedLeftovers = 0;

  while (candles > 0) {
    totalBurned += candles;
    const leftovers = candles + unusedLeftovers;
    candles = Math.floor(leftovers / leftoversNeeded);
    unusedLeftovers = leftovers % leftoversNeeded;
  }

  return totalBurned;
}
```
