---
id: 68216eb60f957572e7c340c4
title: "チャレンジ 11: マイルペース"
challengeType: 28
dashedName: challenge-11
---

# --description--

走ったマイル数と、そのマイルを走るのにかかった時間（`"MM:SS"`（分:秒））が与えられたとき、各マイルを走るのにかかった平均時間を`"MM:SS"`の形式で返してください。

- 必要に応じて先頭にゼロを追加してください。

# --hints--

`milePace(3, "24:00")`は`"08:00"`を返すべきです。

```js
assert.equal(milePace(3, "24:00"), "08:00");
```

`milePace(1, "06:45")`は`"06:45"`を返すべきです。

```js
assert.equal(milePace(1, "06:45"), "06:45");
```

`milePace(2, "07:00")`は`"03:30"`を返すべきです。

```js
assert.equal(milePace(2, "07:00"), "03:30");
```

`milePace(26.2, "120:35")`は`"04:36"`を返すべきです。

```js
assert.equal(milePace(26.2, "120:35"), "04:36");
```

# --seed--

## --seed-contents--

```js
function milePace(miles, duration) {

  return miles;
}
```

# --solutions--

```js
function milePace(miles, duration) {
  const [minutes, seconds] = duration.split(':').map(Number);
  const totalSeconds = minutes * 60 + seconds;
  const avgSecondsPerMile = totalSeconds / miles;

  const avgMinutes = Math.floor(avgSecondsPerMile / 60);
  const avgSeconds = Math.round(avgSecondsPerMile % 60);

  const paddedMinutes = avgMinutes.toString().padStart(2, '0');
  const paddedSeconds = avgSeconds.toString().padStart(2, '0');

  return `${paddedMinutes}:${paddedSeconds}`;
}
```
