---
id: 6821ebda237de8297eaee792
title: "チャレンジ16：丸括弧の反転"
challengeType: 28
dashedName: challenge-16
---

# --description--

正しくネストされた丸括弧を含む文字列が与えられたとき、次のルールに従って文字列のデコード版を返してください。

- 丸括弧の各ペア内のすべての文字は反転されるべきです。
- 丸括弧は最終結果から取り除かれます。
- 丸括弧がネストされている場合、最も内側のペアから先に反転し、その結果を外側のペアの反転に含めます。
- すべての丸括弧は均等にバランスが取れて正しくネストされていると仮定してください。

# --hints--

`decode("(f(b(dc)e)a)")` は `"abcdef"` を返すべきです。

```js
assert.equal(decode("(f(b(dc)e)a)"), "abcdef");
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` は `"Can you read this?"` を返すべきです。

```js
assert.equal(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?");
```

`decode("f(Ce(re))o((e(aC)m)d)p")` は `"freeCodeCamp"` を返すべきです。

```js
assert.equal(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(s) {

  return s;
}
```

# --solutions--

```js
function decode(s) {
  while (s.includes(')')) {
    const closeIndex = s.indexOf(')');
    const openIndex = s.lastIndexOf('(', closeIndex);
    const before = s.slice(0, openIndex);
    const group = s.slice(openIndex + 1, closeIndex).split('').reverse().join('');
    const after = s.slice(closeIndex + 1);
    s = before + group + after;
  }

  return s;
}
```
