---
id: 6821ebe4237de8297eaee794
title: "チャレンジ18: 2番目に高い価格"
challengeType: 28
dashedName: challenge-18
---

# --description--

異なるノートパソコンの価格を表す整数の配列と、予算を表す整数が与えられたとき、次のいずれかを返してください。

1. 予算内であれば2番目に高い価格のノートパソコン、
2. 予算内であれば最も高い価格のノートパソコン、
3. 予算内にノートパソコンがない場合は`0`。

- 重複する価格は無視してください。

# --hints--

`getLaptopCost([1500, 2000, 1800, 1400], 1900)`は`1800`を返すべきです

```js
assert.equal(getLaptopCost([1500, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900)`は`1800`を返すべきです

```js
assert.equal(getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([2099, 1599, 1899, 1499], 2200)`は`1899`を返すべきです

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 2200), 1899);
```

`getLaptopCost([2099, 1599, 1899, 1499], 1000)`は`0`を返すべきです

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 1000), 0);
```

`getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450)`は`1400`を返すべきです

```js
assert.equal(getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450), 1400);
```

# --seed--

## --seed-contents--

```js
function getLaptopCost(laptops, budget) {

  return laptops;
}
```

# --solutions--

```js
function getLaptopCost(laptops, budget) {
  laptops = [...new Set(laptops)].sort((a, b) => b - a);

  if (budget >= laptops[1]) return laptops[1];
  if (budget < laptops[laptops.length - 1]) return 0;

  for (let i = 2; i < laptops.length; i++) {
    if (budget >= laptops[i]) return laptops[i];
  }
}
```
