---
id: 6821ebee237de8297eaee796
title: "チャレンジ 20: 配列の重複"
challengeType: 28
dashedName: challenge-20
---

# --description--

整数の配列が与えられたとき、元の配列に2回以上現れる整数だけを含む配列を、昇順に並べて返してください。2回以上現れる値がなければ、空の配列を返してください。

- 返す配列には各値を1回だけ含めてください。

# --hints--

`findDuplicates([1, 2, 3, 4, 5])` は `[]` を返すべきです。

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 5]), []);
```

`findDuplicates([1, 2, 3, 4, 1, 2])` は `[1, 2]` を返すべきです。

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 1, 2]), [1, 2]);
```

`findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4])` は `[-6, 0, 2, 4, 5, 23]` を返すべきです。

```js
assert.deepEqual(findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4]), [-6, 0, 2, 4, 5, 23]);
```

# --seed--

## --seed-contents--

```js
function findDuplicates(arr) {

  return arr;
}
```

# --solutions--

```js
function findDuplicates(arr) {
  const duplicates = [];
  
  for (let i = 0; i < arr.length - 1; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] === arr[j] && !duplicates.includes(arr[i])) {
        duplicates.push(arr[i]);
      }
    }
  }

  return duplicates.sort((a, b) => a - b);
}
```
