---
id: 6821ebc9237de8297eaee78f
title: "チャレンジ 13: 不自然な素数"
challengeType: 28
dashedName: challenge-13
---

# --description--

整数が与えられたとき、その数が素数か負の素数かを判定してください。

- 素数とは、1より大きい正の整数で、1とその数自身以外に約数を持たない数です。
- 負の素数とは、正の素数の負のバージョンです。
- `1` と `0` は素数とはみなされません。

# --hints--

`isUnnaturalPrime(1)` は `false` を返すべきです。

```js
assert.isFalse(isUnnaturalPrime(1));
```

`isUnnaturalPrime(-1)` は `false` を返すべきです。

```js
assert.isFalse(isUnnaturalPrime(-1));
```

`isUnnaturalPrime(19)` は `true` を返すべきです。

```js
assert.isTrue(isUnnaturalPrime(19));
```

`isUnnaturalPrime(-23)` は `true` を返すべきです。

```js
assert.isTrue(isUnnaturalPrime(-23));
```

`isUnnaturalPrime(0)` は `false` を返すべきです。

```js
assert.isFalse(isUnnaturalPrime(0));
```

`isUnnaturalPrime(97)` は `true` を返すべきです。

```js
assert.isTrue(isUnnaturalPrime(97));
```

`isUnnaturalPrime(-61)` は `true` を返すべきです。

```js
assert.isTrue(isUnnaturalPrime(-61));
```

`isUnnaturalPrime(99)` は `false` を返すべきです。

```js
assert.isFalse(isUnnaturalPrime(99));
```

`isUnnaturalPrime(-44)` は `false` を返すべきです。

```js
assert.isFalse(isUnnaturalPrime(-44));
```

# --seed--

## --seed-contents--

```js
function isUnnaturalPrime(n) {

  return n;
}
```

# --solutions--

```js
function isUnnaturalPrime(n) {
  const abs = Math.abs(n);

  if (abs <= 1) return false;

  for (let i = 2; i <= Math.sqrt(abs); i++) {
    if (abs % i === 0) return false;
  }

  return true;
}
```
