---
id: 68216ef80f957572e7c340c5
title: "チャレンジ 12: メッセージデコーダー"
challengeType: 28
dashedName: challenge-12
---

# --description--

秘密のメッセージ文字列と、そのメッセージをエンコードするために使われた文字数を表す整数が与えられたとき、デコードされた文字列を返してください。

- 正の数はメッセージがアルファベット順に前方へシフトされたことを意味します。
- 負の数はメッセージがアルファベット順に後方へシフトされたことを意味します。
- 大文字・小文字は区別され、デコードされた文字は元のエンコードされた文字の大文字・小文字を保持してください。
- アルファベット以外の文字はデコードされません。

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` は `"This is a secret message."` を返すべきです

```js
assert.equal(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.");
```

`decode("Byffi Qilfx!", 20)` は `"Hello World!"` を返すべきです

```js
assert.equal(decode("Byffi Qilfx!", 20), "Hello World!");
```

`decode("Zqd xnt njzx?", -1)` は `"Are you okay?"` を返すべきです

```js
assert.equal(decode("Zqd xnt njzx?", -1), "Are you okay?");
```

`decode("oannLxmnLjvy", 9)` は `"freeCodeCamp"` を返すべきです

```js
assert.equal(decode("oannLxmnLjvy", 9), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(message, shift) {

  return message;
}
```

# --solutions--

```js
function decode(message, shift) {
  return message.split('').map(char => {
    if (/[a-zA-Z]/.test(char)) {
      const base = char === char.toLowerCase() ? 'a'.charCodeAt(0) : 'A'.charCodeAt(0);
      const charCode = char.charCodeAt(0);
      const offset = (charCode - base - shift + 26) % 26;
      return String.fromCharCode(base + offset);
    } else {
      return char;
    }
  }).join('');
}
```
