---
id: 6821ebf8237de8297eaee798
title: "チャレンジ 22: トリボナッチ数列"
challengeType: 28
dashedName: challenge-22
---

# --description--

トリボナッチ数列は、各数字が直前の3つの数字の合計になる数列です。`0`、`0`、`1`から始めると、数列の最初の10個の数字は`0`、`0`、`1`、`1`、`2`、`4`、`7`、`13`、`24`、`44`です。

トリボナッチ数列の最初の3つの数字を含む配列と、数列の長さを表す整数が与えられたとき、その長さの数列を含む配列を返してください。

- 関数は長さが0以上の任意の数列に対応する必要があります。
- 長さが0の場合は、空の配列を返してください。
- 最初の数字は数列の一部であることに注意してください。

# --hints--

`tribonacciSequence([0, 0, 1], 20)`は`[0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]`を返すべきです。

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 20), [0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]);
```

`tribonacciSequence([21, 32, 43], 1)`は`[21]`を返すべきです。

```js
assert.deepEqual(tribonacciSequence([21, 32, 43], 1), [21]);
```

`tribonacciSequence([0, 0, 1], 0)`は`[]`を返すべきです。

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 0), []);
```

`tribonacciSequence([10, 20, 30], 2)`は`[10, 20]`を返すべきです。

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 2), [10, 20]);
```

`tribonacciSequence([10, 20, 30], 3)`は`[10, 20, 30]`を返すべきです。

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 3), [10, 20, 30]);
```

`tribonacciSequence([123, 456, 789], 8)`は`[123, 456, 789, 1368, 2613, 4770, 8751, 16134]`を返すべきです。

```js
assert.deepEqual(tribonacciSequence([123, 456, 789], 8), [123, 456, 789, 1368, 2613, 4770, 8751, 16134]);
```

# --seed--

## --seed-contents--

```js
function tribonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function tribonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [startSequence[0], startSequence[1]];
  if (length === 3) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2] + + sequence[sequence.length - 3];
    sequence.push(nextValue);
  }
  return sequence;
}
```
