---
id: 6821ebce237de8297eaee790
title: "チャレンジ 14: 文字バトル"
challengeType: 28
dashedName: challenge-14
---

# --description--

あなたの軍隊と敵軍を表す2つの文字列が与えられます。あなたの軍隊の各文字は、敵軍の同じ位置にいる文字と次のルールに従って戦います。

- 文字 `a-z` はそれぞれ `1-26` の強さを持ちます。
- 文字 `A-Z` はそれぞれ `27-52` の強さを持ちます。
- 数字 `0-9` はその数字の値の強さを持ちます。
- その他のすべての文字は強さがゼロです。
- 各文字は1回の戦いにしか参加できません。

各戦いでは、強い文字が勝ちます。勝利数が多い軍隊が戦争に勝ちます。次の値を返してください。

- あなたの軍隊の文字数が敵軍より多い場合は `"Opponent retreated"`。
- 敵軍の文字数があなたの軍隊より多い場合は `"We retreated"`。
- あなたの軍隊がより多くの戦いに勝った場合は `"We won"`。
- 敵軍がより多くの戦いに勝った場合は `"We lost"`。
- 両軍が同じ数の戦いに勝った場合は `"It was a tie"`。

# --hints--

`battle("Hello", "World")` は `"We lost"` を返すべきです。

```js
assert.equal(battle("Hello", "World"), "We lost");
```

`battle("pizza", "salad")` は `"We won"` を返すべきです。

```js
assert.equal(battle("pizza", "salad"), "We won");
```

`battle("C@T5", "D0G$")` は `"We won"` を返すべきです。

```js
assert.equal(battle("C@T5", "D0G$"), "We won");
```

`battle("kn!ght", "orc")` は `"Opponent retreated"` を返すべきです。

```js
assert.equal(battle("kn!ght", "orc"), "Opponent retreated");
```

`battle("PC", "Mac")` は `"We retreated"` を返すべきです。

```js
assert.equal(battle("PC", "Mac"), "We retreated");
```

`battle("Wizards", "Dragons")` は `"It was a tie"` を返すべきです。

```js
assert.equal(battle("Wizards", "Dragons"), "It was a tie");
```

`battle("Mr. Smith", "Dr. Jones")` は `"It was a tie"` を返すべきです。

```js
assert.equal(battle("Mr. Smith", "Dr. Jones"), "It was a tie");
```

# --seed--

## --seed-contents--

```js
function battle(myArmy, opposingArmy) {

  return myArmy;
}
```

# --solutions--

```js
function getStrength(soldier) {
  const soldiers = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ';
  let strength = 0;

  if (/\d/.test(soldier)) {
    strength = parseInt(soldier);
  } else if (soldiers.includes(soldier)) {
    strength = soldiers.indexOf(soldier) + 1;
  }

  return strength;
}

function battle(myArmy, opposingArmy) {
  if (myArmy.length > opposingArmy.length) return 'Opponent retreated';
  if (opposingArmy.length > myArmy.length) return 'We retreated';

  let myWins = 0;
  let theirWins = 0;

  for (let i = 0; i < myArmy.length; i++) {
    const mySoldier = myArmy[i];
    const theirSoldier = opposingArmy[i];

    const myStrength = getStrength(mySoldier);
    const theirStrength = getStrength(theirSoldier);

    if (myStrength > theirStrength) myWins++;
    if (theirStrength > myStrength) theirWins++;
  }

  return myWins > theirWins ? 'We won' : theirWins > myWins ? 'We lost' : 'It was a tie';
}
```
