---
id: 6821ebfd237de8297eaee799
title: "チャレンジ 23: RGB から 16進数へ"
challengeType: 28
dashedName: challenge-23
---

# --description--

CSS の `rgb(r, g, b)` カラー文字列が与えられたとき、その 16進数表現を返してください。

以下は入力に対する出力の例です：

| 入力   | 出力   |
|---------|----------|
| `"rgb(255, 255, 255)"`| `"#ffffff"` |
| `"rgb(1, 2, 3)"` | `"#010203"` |

- すべての文字は小文字にしてください。
- `#` に続けて6文字を返してください。省略形は使わないでください。

# --hints--

`rgbToHex("rgb(255, 255, 255)")` は `"#ffffff"` を返すべきです。

```js
assert.equal(rgbToHex("rgb(255, 255, 255)"), "#ffffff");
```

`rgbToHex("rgb(1, 11, 111)")` は `"#010b6f"` を返すべきです。

```js
assert.equal(rgbToHex("rgb(1, 11, 111)"), "#010b6f");
```

`rgbToHex("rgb(173, 216, 230)")` は `"#add8e6"` を返すべきです。

```js
assert.equal(rgbToHex("rgb(173, 216, 230)"), "#add8e6");
```

`rgbToHex("rgb(79, 123, 201)")` は `"#4f7bc9"` を返すべきです。

```js
assert.equal(rgbToHex("rgb(79, 123, 201)"), "#4f7bc9");
```

# --seed--

## --seed-contents--

```js
function rgbToHex(rgb) {

  return rgb;
}
```

# --solutions--

```js
function rgbToHex(rgb) {
  const match = rgb.match(/\d+/g);
  const [r, g, b] = match.map(num =>
    Math.max(0, Math.min(255, parseInt(num)))
      .toString(16)
      .padStart(2, '0')
  );

  return `#${r}${g}${b}`;
}
```
