---
id: 6821ebdf237de8297eaee793
title: "チャレンジ 17: 演算の順序を無視する"
challengeType: 28
dashedName: challenge-17
---

# --description--

整数の配列と文字列の演算子の配列が与えられたとき、左から右へ順に数値に対して演算を適用してください。すべての数値が使われるまで、必要に応じて演算を繰り返します。最終的な結果を返してください。

例えば、`[1, 2, 3, 4, 5]` と `['+', '*']` が与えられた場合、標準の演算順序を無視して左から右へ `1 + 2 * 3 + 4 * 5` を評価した結果を返します。

- 有効な演算子は `+`、`-`、`*`、`/`、および `%` です。

# --hints--

`evaluate([5, 6, 7, 8, 9], ['+', '-'])` は `3` を返すべきです

```js
assert.equal(evaluate([5, 6, 7, 8, 9], ['+', '-']), 3);
```

`evaluate([17, 61, 40, 24, 38, 14], ['+', '%'])` は `38` を返すべきです

```js
assert.equal(evaluate([17, 61, 40, 24, 38, 14], ['+', '%']), 38);
```

`evaluate([20, 2, 4, 24, 12, 3], ['*', '/'])` は `60` を返すべきです

```js
assert.equal(evaluate([20, 2, 4, 24, 12, 3], ['*', '/']), 60);
```

`evaluate([11, 4, 10, 17, 2], ['*', '*', '%'])` は `30` を返すべきです

```js
assert.equal(evaluate([11, 4, 10, 17, 2], ['*', '*', '%']), 30);
```

`evaluate([33, 11, 29, 13], ['/', '-'])` は `-2` を返すべきです

```js
assert.equal(evaluate([33, 11, 29, 13], ['/', '-']), -2);
```

# --seed--

## --seed-contents--

```js
function evaluate(numbers, operators) {

  return numbers;
}
```

# --solutions--

```js
function doMath(a, b, operator) {
  switch (operator) {
    case "+":
      return a + b;
    case "-":
      return a - b;
    case "*":
      return a * b;
    case "/":
      return a / b;
    default:
      return a % b;
  }
}

function evaluate(numbers, operators) { 
  let total = numbers[0];

  for (let i = 1; i < numbers.length; i++) {
    const operator = operators[(i - 1) % operators.length];
    total = doMath(total, numbers[i], operator);
  }

  return total;
}
```
