---
id: 68adce01c0e1144d0a902958
title: "チャレンジ 26: IPv4 バリデーター"
challengeType: 28
dashedName: challenge-26
---

# --description--

文字列が有効なIPv4アドレスかどうかを判定してください。有効なIPv4アドレスは、ドット（`.`）で区切られた4つの整数で構成されています。各数字は以下の条件を満たす必要があります。

- 0から255までの範囲内であること（両端を含む）。
- 先頭にゼロがないこと（例：0は許可されますが、01は許可されません）。
- 数字のみが許可されます。

# --hints--

`isValidIPv4("192.168.1.1")` は `true` を返すべきです。

```js
assert.isTrue(isValidIPv4("192.168.1.1"));
```

`isValidIPv4("0.0.0.0")` は `true` を返すべきです。

```js
assert.isTrue(isValidIPv4("0.0.0.0"));
```

`isValidIPv4("255.255.255.255")` は `true` を返す必要があります。

```js
assert.isTrue(isValidIPv4("255.255.255.255"));
```

`isValidIPv4("208.67.222.222")` は `true` を返す必要があります。

```js
assert.isTrue(isValidIPv4("208.67.222.222"));
```

`isValidIPv4("255.01.50.111")` は `false` を返すべきです。

```js
assert.isFalse(isValidIPv4("255.01.50.111"));
```

`isValidIPv4("255.00.50.111")` は `false` を返すべきです。

```js
assert.isFalse(isValidIPv4("255.00.50.111"));
```

`isValidIPv4("256.101.50.115")` は `false` を返すべきです。

```js
assert.isFalse(isValidIPv4("256.101.50.115"));
```

`isValidIPv4("192.168.101.")` は `false` を返すべきです。

```js
assert.isFalse(isValidIPv4("192.168.101."));
```

`isValidIPv4("192168145213")` は `false` を返すべきです。

```js
assert.isFalse(isValidIPv4("192168145213"));
```

`isValidIPv4("192.168.1.a")` は `false` を返す必要があります。

```js
assert.isFalse(isValidIPv4("192.168.1.a"));
```

`isValidIPv4("192.168.o.1")` は `false` を返す必要があります。

```js
assert.isFalse(isValidIPv4("192.168.o.1"));
```

`isValidIPv4("192.168.1.1a")` は `false` を返す必要があります。

```js
assert.isFalse(isValidIPv4("192.168.1.1a"));
```

# --seed--

## --seed-contents--

```js
function isValidIPv4(ipv4) {

  return ipv4;
}
```

# --solutions--

```js
function isValidIPv4(ipv4) {
  const parts = ipv4.split(".");
  
  if (parts.length !== 4) return false;

  for (let part of parts) {
    if (!/^\d+$/.test(part)) return false;    

    const num = Number(part);
    if (num < 0 || num > 255) return false;

    if (part.length > 1 && part.startsWith("0")) return false;
  }

  return true;
}
```
