---
id: 68adce01c0e1144d0a90295c
title: "チャレンジ 28: ローマ数字パーサー"
challengeType: 28
dashedName: challenge-28
---

# --description--

ローマ数字を表す文字列が与えられたら、その整数値を返してください。

ローマ数字は以下の記号と値で構成されています：

| 記号 | 値   |
|--------|-------|
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

- 数字は左から右に読みます。小さい数字が大きい数字の前に現れた場合、その値は引かれます。それ以外の場合は値が加算されます。

# --hints--

`parseRomanNumeral("III")` は `3` を返すべきです。

```js
assert.equal(parseRomanNumeral("III"), 3);
```

`parseRomanNumeral("IV")` は `4` を返すべきです。

```js
assert.equal(parseRomanNumeral("IV"), 4);
```

`parseRomanNumeral("XXVI")` は `26` を返すべきです。

```js
assert.equal(parseRomanNumeral("XXVI"), 26);
```

`parseRomanNumeral("XCIX")` は `99` を返すべきです。

```js
assert.equal(parseRomanNumeral("XCIX"), 99);
```

`parseRomanNumeral("CDLX")` は `460` を返すべきです。

```js
assert.equal(parseRomanNumeral("CDLX"), 460);
```

`parseRomanNumeral("DIV")` は `504` を返すべきです。

```js
assert.equal(parseRomanNumeral("DIV"), 504);
```

`parseRomanNumeral("MMXXV")` は `2025` を返すべきです。

```js
assert.equal(parseRomanNumeral("MMXXV"), 2025);
```

# --seed--

## --seed-contents--

```js
function parseRomanNumeral(numeral) {

  return numeral;
}
```

# --solutions--

```js
function parseRomanNumeral(numeral) {
  const romanMap = {
    I: 1,
    V: 5,
    X: 10,
    L: 50,
    C: 100,
    D: 500,
    M: 1000
  };

  let total = 0;

  for (let i = 0; i < numeral.length; i++) {
    const current = romanMap[numeral[i]];
    const next = romanMap[numeral[i + 1]];

    if (next && current < next) {
      total -= current;
    } else {
      total += current;
    }
  }

  return total;
}
```
