---
id: 68adce01c0e1144d0a90295a
title: "チャレンジ 27: 行列の回転"
challengeType: 28
dashedName: challenge-27
---

# --description--

行列（配列の配列）が与えられたとき、その行列を90度時計回りに回転させて返してください。例えば、`[[1, 2], [3, 4]]` は次のようになっています：

| 1 | 2 |
|---|---|
| 3 | 4 |

これに対して、`[[3, 1], [4, 2]]` は次のようになります：

| 3 | 1 |
|---|---|
| 4 | 2 |

# --hints--

`rotate([[1]])` は `[[1]]` を返すべきです。

```js
assert.deepEqual(rotate([[1]]), [[1]]);
```

`rotate([[1, 2], [3, 4]])` は `[[3, 1], [4, 2]]` を返すべきです。

```js
assert.deepEqual(rotate([[1, 2], [3, 4]]), [[3, 1], [4, 2]]);
```

`rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]])` は `[[7, 4, 1], [8, 5, 2], [9, 6, 3]]` を返すべきです。

```js
assert.deepEqual(rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]]), [[7, 4, 1], [8, 5, 2], [9, 6, 3]]);
```

`rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]])` は `[[0, 1, 0], [0, 0, 1], [0, 1, 0]]` を返すべきです。

```js
assert.deepEqual(rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]]), [[0, 1, 0], [0, 0, 1], [0, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function rotate(matrix) {

  return matrix;
}
```

# --solutions--

```js
function rotate(matrix) {
  const n = matrix.length;
  const result = Array.from({ length: n }, () => Array(n).fill(0));

  for (let i = 0; i < n; i++) {
    for (let j = 0; j < n; j++) {
      result[j][n - 1 - i] = matrix[i][j];
    }
  }

  return result;
}
```
