---
id: 68adce01c0e1144d0a902956
title: "チャレンジ 25: 母音リピーター"
challengeType: 28
dashedName: challenge-25
---

# --description--

文字列が与えられたとき、各母音が前に出てきた母音よりも1回多く繰り返される新しい文字列を返してください。たとえば、文中の最初の母音はそのままです。2番目の母音は連続して2回、3番目の母音は連続して3回、というように続きます。

- `a`、`e`、`i`、`o`、`u` の文字は、大文字・小文字にかかわらず母音とみなします。
- 元の母音は元の大文字・小文字のままにしてください。
- 繰り返す母音は小文字にしてください。
- 母音以外の文字は元の大文字・小文字のままにしてください。

# --hints--

`repeatVowels("hello world")` は `"helloo wooorld"` を返すべきです。

```js
assert.equal(repeatVowels("hello world"), "helloo wooorld");
```

`repeatVowels("freeCodeCamp")` は `"freeeCooodeeeeCaaaaamp"` を返すべきです。

```js
assert.equal(repeatVowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp");
```

`repeatVowels("AEIOU")` は `"AEeIiiOoooUuuuu"` を返すべきです。

```js
assert.equal(repeatVowels("AEIOU"), "AEeIiiOoooUuuuu");
```

`repeatVowels("I like eating ice cream in Iceland")` は `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"` を返すべきです。

```js
assert.equal(repeatVowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand");
```

# --seed--

## --seed-contents--

```js
function repeatVowels(str) {

  return str;
}
```

# --solutions--

```js
function repeatVowels(str) {
  const vowels = "aeiouAEIOU";
  let count = 0;
  let result = "";

  for (let char of str) {
    result += char;
    if (vowels.includes(char)) {
      result += char.repeat(count).toLowerCase();
      count++;
    }
  }

  return result;
}
```
