---
id: 68af0687ef34c76c28ffa54b
title: "チャレンジ 32: 文の逆順"
challengeType: 28
dashedName: challenge-32
---

# --description--

単語の文字列が与えられたら、単語の順序を逆にした新しい文字列を返してください。例えば、最初の単語は返される文字列の最後に、最後の単語は返される文字列の最初に来るようにします。

- 与えられた文字列では、単語は1つ以上のスペースで区切られている場合があります。
- 返される文字列では、単語間は1つのスペースだけにしてください。

# --hints--

`reverseSentence("world hello")` は `"hello world"` を返すべきです。

```js
assert.equal(reverseSentence("world hello"), "hello world");
```

`reverseSentence("push commit git")` は `"git commit push"` を返すべきです。

```js
assert.equal(reverseSentence("push commit git"), "git commit push");
```

`reverseSentence("npm  install  sudo")` は `"sudo install npm"` を返すべきです。

```js
assert.equal(reverseSentence("npm  install   apt    sudo"), "sudo apt install npm");
```

`reverseSentence("import    default   function  export")` は `"export function default import"` を返すべきです。

```js
assert.equal(reverseSentence("import    default   function  export"), "export function default import");
```

# --seed--

## --seed-contents--

```js
function reverseSentence(sentence) {

  return sentence;
}
```

# --solutions--

```js
function reverseSentence(sentence) {
  return sentence.split(/\s+/).reverse().join(" ");
}
```
