---
id: 68af0687ef34c76c28ffa54f
title: "チャレンジ 34: 欠けている数字"
challengeType: 28
dashedName: challenge-34
---

# --description--

1から`n`までの整数が含まれる配列が与えられたとき、1から`n`までの間で（`n`が与えられた配列の最大の数字です）欠けているすべての整数を配列で返してください。

- 与えられた配列はソートされていない場合や重複が含まれている場合があります。
- 返す配列は昇順である必要があります。
- 欠けている整数がなければ、空の配列を返してください。

# --hints--

`findMissingNumbers([1, 3, 5])`は`[2, 4]`を返すべきです。

```js
assert.deepEqual(findMissingNumbers([1, 3, 5]), [2, 4]);
```

`findMissingNumbers([1, 2, 3, 4, 5])`は`[]`を返すべきです。

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5]), []);
```

`findMissingNumbers([1, 10])`は`[2, 3, 4, 5, 6, 7, 8, 9]`を返すべきです。

```js
assert.deepEqual(findMissingNumbers([1, 10]), [2, 3, 4, 5, 6, 7, 8, 9]);
```

`findMissingNumbers([10, 1, 10, 1, 10, 1])`は`[2, 3, 4, 5, 6, 7, 8, 9]`を返すべきです。

```js
assert.deepEqual(findMissingNumbers([10, 1, 10, 1, 10, 1]), [2, 3, 4, 5, 6, 7, 8, 9]);
```

`findMissingNumbers([3, 1, 4, 1, 5, 9])`は`[2, 6, 7, 8]`を返すべきです。

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5]), []);
```

`findMissingNumbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4])`は`[11]`を返すべきです。

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4]), [11]);
```

# --seed--

## --seed-contents--

```js
function findMissingNumbers(arr) {

  return arr;
}
```

# --solutions--

```js
function findMissingNumbers(arr) {
  if (arr.length === 0) return [];
  
  const n = Math.max(...arr);
  const set = new Set(arr);
  const result = [];

  for (let i = 1; i <= n; i++) {
    if (!set.has(i)) {
      result.push(i);
    }
  }
  return result;
}
```
