---
id: 68af0687ef34c76c28ffa547
title: "チャレンジ 30: ユニークな文字"
challengeType: 28
dashedName: challenge-30
---

# --description--

文字列が与えられたとき、その文字列内のすべての文字がユニークかどうかを判定してください。

- 大文字と小文字は異なる文字として扱います。

# --hints--

`allUnique("abc")` は `true` を返すべきです。

```js
assert.isTrue(allUnique("abc"));
```

`allUnique("aA")` は `true` を返すべきです。

```js
assert.isTrue(allUnique("aA"));
```

`allUnique("QwErTy123!@")` は `true` を返すべきです。

```js
assert.isTrue(allUnique("QwErTy123!@"));
```

`allUnique("~!@#$%^&*()_+")` は `true` を返すべきです。

```js
assert.isTrue(allUnique("~!@#$%^&*()_+"));
```

`allUnique("hello")` は `false` を返すべきです。

```js
assert.isFalse(allUnique("hello"));
```

`allUnique("freeCodeCamp")` は `false` を返すべきです。

```js
assert.isFalse(allUnique("freeCodeCamp"));
```

`allUnique("!@#*$%^&*()aA")` は `false` を返すべきです。

```js
assert.isFalse(allUnique("!@#*$%^&*()aA"));
```

# --seed--

## --seed-contents--

```js
function allUnique(str) {

  return str;
}
```

# --solutions--

```js
function allUnique(str) {
  let seen = '';
  for (let char of str) {
    if (seen.includes(char)) {
      return false;
    }
    seen += char;
  }
  return true;
}
```
