---
id: 68af0687ef34c76c28ffa54d
title: "チャレンジ 33: スクリーンタイム"
challengeType: 28
dashedName: challenge-33
---

# --description--

7つの整数からなる入力配列が与えられます。これは1週間の時間を表しており、それぞれの整数はその日にスマホを使った時間（時間数）です。以下の条件に基づいて、スクリーンタイムが多すぎるかどうかを判定してください。

- どの日か1日でも10時間以上あれば、多すぎます。
- 連続する3日間の平均が8時間以上であれば、多すぎます。
- 7日間の平均が6時間以上であれば、多すぎます。

# --hints--

`tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7])` は `false` を返すべきです。

```js
assert.isFalse(tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7]));
```

`tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3])` は `false` を返すべきです。

```js
assert.isFalse(tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3]));
```

`tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6])` は `false` を返すべきです。

```js
assert.isFalse(tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6]));
```

`tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4])` は `true` を返すべきです。

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4]));
```

`tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0])` は `true` を返すべきです。

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0]));
```

`tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4])` は `true` を返すべきです。

```js
assert.isTrue(tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4]));
```

`tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6])` は `true` を返すべきです。

```js
assert.isTrue(tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6]));
```

# --seed--

## --seed-contents--

```js
function tooMuchScreenTime(hours) {

  return hours;
}
```

# --solutions--

```js
function tooMuchScreenTime(hours) {
  for (let h of hours) {
    if (h >= 10) return true;
  }

  for (let i = 0; i <= 4; i++) {
    const avg3 = (hours[i] + hours[i+1] + hours[i+2]) / 3;
    if (avg3 >= 8) return true;
  }

  const weeklyAvg = hours.reduce((sum, h) => sum + h, 0) / 7;
  if (weeklyAvg >= 6) return true;

  return false;
}
```
