---
id: 68b06e589bf2273243814771
title: "チャレンジ 36: サーモスタット調整器"
challengeType: 28
dashedName: challenge-36
---

# --description--

部屋の現在の温度と目標温度が与えられたとき、次の条件に基づいて部屋の温度をどのように調整するかを示す文字列を返してください。

- 現在の温度が目標温度より低い場合は `"heat"` を返してください。
- 現在の温度が目標温度より高い場合は `"cool"` を返してください。
- 現在の温度が目標温度と等しい場合は `"hold"` を返してください。

# --hints--

`adjustThermostat(68, 72)` は `"heat"` を返すべきです。

```js
assert.equal(adjustThermostat(68, 72), "heat");
```

`adjustThermostat(75, 72)` は `"cool"` を返すべきです。

```js
assert.equal(adjustThermostat(75, 72), "cool");
```

`adjustThermostat(72, 72)` は `"hold"` を返すべきです。

```js
assert.equal(adjustThermostat(72, 72), "hold");
```

`adjustThermostat(-20.5, -10.1)` は `"heat"` を返すべきです。

```js
assert.equal(adjustThermostat(-20.5, -10.1), "heat");
```

`adjustThermostat(100, 99.9)` は `"cool"` を返すべきです。

```js
assert.equal(adjustThermostat(100, 99.9), "cool");
```

`adjustThermostat(0.0, 0.0)` は `"hold"` を返すべきです。

```js
assert.equal(adjustThermostat(0.0, 0.0), "hold");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(temp, target) {

  return temp;
}
```

# --solutions--

```js
function adjustThermostat(temp, target) {
  if (temp < target) {
    return "heat";
  } else if (temp > target) {
    return "cool";
  } else {
    return "hold";
  }
}
```
