---
id: 68b06e589bf2273243814777
title: "チャレンジ 39: タンクを満たす"
challengeType: 28
dashedName: challenge-39
---

# --description--

燃料タンクの容量、現在の燃料量、ガロンあたりの価格が与えられたとき、タンクを満タンにするための費用を返してください。

- `tankSize` はタンクの総容量（ガロン単位）です。
- `fuelLevel` はタンク内の現在の燃料量（ガロン単位）です。
- `pricePerGallon` は1ガロンあたりの燃料の価格です。
- 返す値は小数点以下2桁に丸め、`"$d.dd"` の形式で返してください。

# --hints--

`costToFill(20, 0, 4.00)` は `"$80.00"` を返すべきです。

```js
assert.equal(costToFill(20, 0, 4.00), "$80.00");
```

`costToFill(15, 10, 3.50)` は `"$17.50"` を返すべきです。

```js
assert.equal(costToFill(15, 10, 3.50), "$17.50");
```

`costToFill(18, 9, 3.25)` は `"$29.25"` を返すべきです。

```js
assert.equal(costToFill(18, 9, 3.25), "$29.25");
```

`costToFill(12, 12, 4.99)` は `"$0.00"` を返すべきです。

```js
assert.equal(costToFill(12, 12, 4.99), "$0.00");
```

`costToFill(15, 9.5, 3.98)` は `"$21.89"` を返すべきです。

```js
assert.equal(costToFill(15, 9.5, 3.98), "$21.89");
```

# --seed--

## --seed-contents--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {

  return tankSize;
}
```

# --solutions--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {
  let gallonsNeeded = tankSize - fuelLevel;
  let cost = gallonsNeeded * pricePerGallon;
  return `$${cost.toFixed(2)}`;
}
```
