---
id: 68b06e589bf2273243814775
title: "チャレンジ 38: スラッグジェネレーター"
challengeType: 28
dashedName: challenge-38
---

# --description--

文字列が与えられたら、次の条件を満たすURLに適した文字列を返してください。

- すべての文字は小文字にしてください。
- 文字、数字、スペース以外のすべての文字は削除してください。
- すべてのスペースはURLエンコードされたスペースコード`%20`に置き換えてください。
- 連続するスペースは1つの`%20`に置き換えてください。
- 返される文字列の先頭と末尾に`%20`があってはいけません。

# --hints--

`generateSlug("helloWorld")`は`"helloworld"`を返すはずです。

```js
assert.equal(generateSlug("helloWorld"), "helloworld");
```

`generateSlug("hello world!")`は`"hello%20world"`を返すはずです。

```js
assert.equal(generateSlug("hello world!"), "hello%20world");
```

`generateSlug(" hello-world ")`は`"helloworld"`を返すはずです。

```js
assert.equal(generateSlug(" hello-world "), "helloworld");
```

`generateSlug("hello  world")`は`"hello%20world"`を返すはずです。

```js
assert.equal(generateSlug("hello  world"), "hello%20world");
```

`generateSlug("  ?H^3-1*1]0! W[0%R#1]D  ")`は`"h3110%20w0r1d"`を返すはずです。

```js
assert.equal(generateSlug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d");
```

# --seed--

## --seed-contents--

```js
function generateSlug(str) {

  return str;
}
```

# --solutions--

```js
function generateSlug(str) {
  let cleaned = str.replace(/[^a-zA-Z0-9 ]+/g, "");
  cleaned = cleaned.replace(/\s+/g, " ").trim();
  return cleaned.toLowerCase().replace(/ /g, "%20");
}
```
