---
id: 68b06e589bf227324381476f
title: "チャレンジ 35: 単語の出現頻度"
challengeType: 28
dashedName: challenge-35
---

# --description--

段落が与えられたとき、最も頻繁に出現する単語を3つ返す配列を作成してください。

- 段落内の単語はスペースで区切られています。
- 段落内の大文字・小文字は区別しません。例えば、`Hello` と `hello` は同じ単語として扱います。
- 段落内の句読点は無視してください。句読点にはカンマ（`,`）、ピリオド（`.`）、感嘆符（`!`）が含まれます。
- 返す配列の単語はすべて小文字にしてください。
- 返す配列は、最も頻繁に出現する単語が最初に来るように降順で並べてください。

# --hints--

`getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` は `["coding", "python", "in"]` を返すべきです。

```js
assert.deepEqual(getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"]);
```

`getWords("I like coding. I like testing. I love debugging!")` は `["i", "like", "coding"]` を返すべきです。

```js
assert.deepEqual(getWords("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"]);
```

`getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` は `["debug", "test", "deploy"]` を返すべきです。

```js
assert.deepEqual(getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"]);
```

# --seed--

## --seed-contents--

```js
function getWords(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function getWords(paragraph) {
  const cleaned = paragraph.replace(/[.,!]/g, "").toLowerCase();
  const words = cleaned.split(/\s+/);

  const freq = {};
  for (const word of words) {
    if (word) {
      freq[word] = (freq[word] || 0) + 1;
    }
  }

  const sorted = Object.keys(freq).sort((a, b) => freq[b] - freq[a]);
  return sorted.slice(0, 3);
}
```
