---
id: 68b06e589bf2273243814773
title: "チャレンジ 37: 文の先頭を大文字にする"
challengeType: 28
dashedName: challenge-37
---

# --description--

段落が与えられたとき、各文の最初の文字を大文字にした新しい段落を返してください。

- その他の文字はすべてそのまま保持してください。
- 文の終わりはピリオド (`.`)、1つ以上の疑問符 (`?`)、または1つ以上の感嘆符 (`!`) で終わる場合があります。

# --hints--

`capitalize("this is a simple sentence.")` は `"This is a simple sentence."` を返すべきです。

```js
assert.equal(capitalize("this is a simple sentence."), "This is a simple sentence.");
```

`capitalize("hello world. how are you?")` は `"Hello world. How are you?"` を返すべきです。

```js
assert.equal(capitalize("hello world. how are you?"), "Hello world. How are you?");
```

`capitalize("i did today's coding challenge... it was fun!!")` は `"I did today's coding challenge... It was fun!!"` を返すべきです。

```js
assert.equal(capitalize("i did today's coding challenge... it was fun!!"), "I did today's coding challenge... It was fun!!");
```

`capitalize("crazy!!!strange???unconventional...sentences.")` は `"Crazy!!!Strange???Unconventional...Sentences."` を返すべきです。

```js
assert.equal(capitalize("crazy!!!strange???unconventional...sentences."), "Crazy!!!Strange???Unconventional...Sentences.");
```

`capitalize("there's a space before this period . why is there a space before that period ?")` は `"There's a space before this period . Why is there a space before that period ?"` を返すべきです。

```js
assert.equal(capitalize("there's a space before this period . why is there a space before that period ?"), "There's a space before this period . Why is there a space before that period ?");
```

# --seed--

## --seed-contents--

```js
function capitalize(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function capitalize(paragraph) {
  return paragraph.replace(/([.!?]+)(\s*)([a-z])/g, (match, punc, spaces, char) => {
    return punc + spaces + char.toUpperCase();
  })
  .replace(/^([a-z])/, (match, char) => char.toUpperCase());
}
```
