---
id: 68b1f72371a5ac895ac70a02
title: "チャレンジ 40: 写真の保存"
challengeType: 28
dashedName: challenge-40
---

# --description--

写真のサイズ（メガバイト単位）とハードドライブの容量（ギガバイト単位）が与えられたとき、以下の条件に従ってハードドライブに保存できる写真の枚数を返してください。

- 1ギガバイトは1000メガバイトに相当します。
- ハードドライブに保存できる写真の枚数は整数で返してください。

# --hints--

`numberOfPhotos(1, 1)` は `1000` を返すべきです。

```js
assert.equal(numberOfPhotos(1, 1), 1000);
```

`numberOfPhotos(2, 1)` は `500` を返すべきです。

```js
assert.equal(numberOfPhotos(2, 1), 500);
```

`numberOfPhotos(4, 256)` は `64000` を返すべきです。

```js
assert.equal(numberOfPhotos(4, 256), 64000);
```

`numberOfPhotos(3.5, 750)` は `214285` を返すべきです。

```js
assert.equal(numberOfPhotos(3.5, 750), 214285);
```

`numberOfPhotos(3.5, 5.5)` は `1571` を返すべきです。

```js
assert.equal(numberOfPhotos(3.5, 5.5), 1571);
```

# --seed--

## --seed-contents--

```js
function numberOfPhotos(photoSizeMb, hardDriveSizeGb) {

  return photoSizeMb;
}
```

# --solutions--

```js
function numberOfPhotos(photoSizeMb, driveSizeGb) {
  const driveSizeMb = driveSizeGb * 1000;
  return Math.floor(driveSizeMb / photoSizeMb);
}
```
