---
id: 68b1f72371a5ac895ac70a08
title: "チャレンジ 43: 数字と文字の比較"
challengeType: 28
dashedName: challenge-43
---

# --description--

文字列が与えられたとき、その文字列に数字が文字より多ければ `"digits"` を、文字が数字より多ければ `"letters"` を、数字と文字の数が同じであれば `"tie"` を返してください。

- 数字は `0-9` で構成されます。
- 文字は大文字または小文字の `a-z` で構成されます。
- その他の文字は無視してください。

# --hints--

`digitsOrLetters("abc123")` は `"tie"` を返すべきです。

```js
assert.equal(digitsOrLetters("abc123"), "tie");
```

`digitsOrLetters("a1b2c3d")` は `"letters"` を返すべきです。

```js
assert.equal(digitsOrLetters("a1b2c3d"), "letters");
```

`digitsOrLetters("1a2b3c4")` は `"digits"` を返すべきです。

```js
assert.equal(digitsOrLetters("1a2b3c4"), "digits");
```

`digitsOrLetters("abc123!@#DEF")` は `"letters"` を返すべきです。

```js
assert.equal(digitsOrLetters("abc123!@#DEF"), "letters");
```

`digitsOrLetters("H3110 W0R1D")` は `"digits"` を返すべきです。

```js
assert.equal(digitsOrLetters("H3110 W0R1D"), "digits");
```

`digitsOrLetters("P455W0RD")` は `"tie"` を返すべきです。

```js
assert.equal(digitsOrLetters("P455W0RD"), "tie");
```

# --seed--

## --seed-contents--

```js
function digitsOrLetters(str) {

  return str;
}
```

# --solutions--

```js
function digitsOrLetters(str) {
  let digitCount = 0;
  let letterCount = 0;

  for (let char of str) {
    if (/[0-9]/.test(char)) digitCount++;
    else if (/[a-zA-Z]/.test(char)) letterCount++;
  }

  if (digitCount > letterCount) return "digits";
  if (letterCount > digitCount) return "letters";
  return "tie";
}
```
