---
id: 68b1f72371a5ac895ac70a06
title: "チャレンジ 42: ビデオストレージ"
challengeType: 28
dashedName: challenge-42
---

# --description--

ビデオのサイズ、その単位、ハードドライブの容量、その単位が与えられたとき、以下の条件に従ってハードドライブに保存できるビデオの数を返してください。

- ビデオサイズの単位はバイト（`"B"`）、キロバイト（`"KB"`）、メガバイト（`"MB"`）、ギガバイト（`"GB"`）のいずれかです。
- 上記のビデオ単位でない場合は`"Invalid video unit"`を返してください。
- ハードドライブ容量の単位はギガバイト（`"GB"`）またはテラバイト（`"TB"`）のいずれかです。
- 上記のハードドライブ単位でない場合は`"Invalid drive unit"`を返してください。
- ハードドライブに収まるビデオの整数個数を返してください。
- 以下の換算を使用してください：

| 単位 | 等価値     |
|:----:|:----------:|
| 1 B  | 1 B        |
| 1 KB | 1000 B     |
| 1 MB | 1000 KB    |
| 1 GB | 1000 MB    |
| 1 TB | 1000 GB    |

例えば、`500`、`"MB"`、`100`、`"GB"`を引数として与えられた場合、500 MBのビデオが100 GBのハードドライブにいくつ収まるかを求めてください。

# --hints--

`numberOfVideos(500, "MB", 100, "GB")`は`200`を返すべきです。

```js
assert.equal(numberOfVideos(500, "MB", 100, "GB"), 200);
```

`numberOfVideos(1, "TB", 10, "TB")`は`"Invalid video unit"`を返すべきです。

```js
assert.equal(numberOfVideos(1, "TB", 10, "TB"), "Invalid video unit");
```

`numberOfVideos(2000, "MB", 100000, "MB")`は`"Invalid drive unit"`を返すべきです。

```js
assert.equal(numberOfVideos(2000, "MB", 100000, "MB"), "Invalid drive unit");
```

`numberOfVideos(500000, "KB", 2, "TB")`は`4000`を返すべきです。

```js
assert.equal(numberOfVideos(500000, "KB", 2, "TB"), 4000);
```

`numberOfVideos(1.5, "GB", 2.2, "TB")`は`1466`を返すべきです。

```js
assert.equal(numberOfVideos(1.5, "GB", 2.2, "TB"), 1466);
```

# --seed--

## --seed-contents--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {

  return videoSize;
}
```

# --solutions--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {
  const videoUnits = { "KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000 };
  const driveUnits = { "GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000 };

  if (!videoUnits[videoUnit]) return "Invalid video unit";
  if (!driveUnits[driveUnit]) return "Invalid drive unit";

  const videoBytes = videoSize * videoUnits[videoUnit];
  const driveBytes = driveSize * driveUnits[driveUnit];

  return Math.floor(driveBytes / videoBytes);
}
```
