---
id: 68b1f72371a5ac895ac70a04
title: "チャレンジ 41: ファイルストレージ"
challengeType: 28
dashedName: challenge-41
---

# --description--

ファイルサイズ、ファイルサイズの単位、およびギガバイト（GB）単位のハードドライブ容量が与えられたとき、以下の条件を使ってハードドライブに保存できるファイル数を返してください。

- ファイルサイズの単位はバイト（`"B"`）、キロバイト（`"KB"`）、またはメガバイト（`"MB"`）のいずれかです。
- ドライブに収まるファイルの整数個数を返してください。
- 以下の換算を使用してください：

| 単位 | 同等の量   |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |

たとえば、`500`、`"KB"`、`1`を引数として与えた場合、1 GBのハードドライブに500 KBのファイルがいくつ収まるかを求めてください。

# --hints--

`numberOfFiles(500, "KB", 1)` は `2000` を返すべきです。

```js
assert.equal(numberOfFiles(500, "KB", 1), 2000);
```

`numberOfFiles(50000, "B", 1)` は `20000` を返すべきです。

```js
assert.equal(numberOfFiles(50000, "B", 1), 20000);
```

`numberOfFiles(5, "MB", 1)` は `200` を返すべきです。

```js
assert.equal(numberOfFiles(5, "MB", 1), 200);
```

`numberOfFiles(4096, "B", 1.5)` は `366210` を返すべきです。

```js
assert.equal(numberOfFiles(4096, "B", 1.5), 366210);
```

`numberOfFiles(220.5, "KB", 100)` は `453514` を返すべきです。

```js
assert.equal(numberOfFiles(220.5, "KB", 100), 453514);
```

`numberOfFiles(4.5, "MB", 750)` は `166666` を返すべきです。

```js
assert.equal(numberOfFiles(4.5, "MB", 750), 166666);
```

# --seed--

## --seed-contents--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {

  return fileSize;
}
```

# --solutions--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {
  const driveSizeBytes = driveSizeGb * 1000 * 1000 * 1000;

  let fileSizeBytes;
  if (fileUnit === "B") {
    fileSizeBytes = fileSize;
  } else if (fileUnit === "KB") {
    fileSizeBytes = fileSize * 1000;
  } else {
    fileSizeBytes = fileSize * 1000 * 1000;
  }

  return Math.floor(driveSizeBytes / fileSizeBytes);
}
```
