---
id: 68b1f72371a5ac895ac70a0a
title: "チャレンジ 44: 文字列ミラー"
challengeType: 28
dashedName: challenge-44
---

# --description--

2つの文字列が与えられたとき、2番目の文字列が1番目の文字列のミラー（鏡像）かどうかを判定してください。

- 文字列がミラーとみなされるのは、同じ文字が逆順で含まれている場合です。
- 大文字と小文字は区別してください。
- アルファベット以外の文字は無視してください。

# --hints--

`isMirror("helloworld", "helloworld")` は `false` を返すべきです。

```js
assert.isFalse(isMirror("helloworld", "helloworld"));
```

`isMirror("Hello World", "dlroW olleH")` は `true` を返すべきです。

```js
assert.isTrue(isMirror("Hello World", "dlroW olleH"));
```

`isMirror("RaceCar", "raCecaR")` は `true` を返すべきです。

```js
assert.isTrue(isMirror("RaceCar", "raCecaR"));
```

`isMirror("RaceCar", "RaceCar")` は `false` を返すべきです。

```js
assert.isFalse(isMirror("RaceCar", "RaceCar"));
```

`isMirror("Mirror", "rorrim")` は `false` を返すべきです。

```js
assert.isFalse(isMirror("Mirror", "rorrim"));
```

`isMirror("Hello World", "dlroW-olleH")` は `true` を返すべきです。

```js
assert.isTrue(isMirror("Hello World", "dlroW-olleH"));
```

`isMirror("Hello World", "!dlroW !olleH")` は `true` を返すべきです。

```js
assert.isTrue(isMirror("Hello World", "!dlroW !olleH"));
```

# --seed--

## --seed-contents--

```js
function isMirror(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirror(str1, str2) {
  const clean1 = str1.replace(/[^a-zA-Z]/g, "");
  const clean2 = str2.replace(/[^a-zA-Z]/g, "");
  return clean1.split("").reverse().join("") === clean2;
}
```
