---
id: 68b7687dded630607aceccb3
title: "チャレンジ 49: CSV ヘッダーパーサー"
challengeType: 28
dashedName: challenge-49
---

# --description--

カンマ区切り値（CSV）ファイルの最初の行が与えられたとき、見出しを含む配列を返してください。

- CSVファイルの最初の行には、カンマで区切られた見出しが含まれています。
- 各見出しの前後の空白はすべて取り除いてください。

# --hints--

`getHeadings("name,age,city")` は `["name", "age", "city"]` を返す必要があります。

```js
assert.deepEqual(getHeadings("name,age,city"), ["name", "age", "city"]);
```

`getHeadings("first name,last name,phone")` は `["first name", "last name", "phone"]` を返す必要があります。

```js
assert.deepEqual(getHeadings("first name,last name,phone"), ["first name", "last name", "phone"]);
```

`getHeadings("username , email , signup date ")` は `["username", "email", "signup date"]` を返す必要があります。

```js
assert.deepEqual(getHeadings("username , email , signup date "), ["username", "email", "signup date"]);
```

# --seed--

## --seed-contents--

```js
function getHeadings(csv) {

  return csv;
}
```

# --solutions--

```js
function getHeadings(csv) {
  return csv
    .split(",")
    .map(h => h.trim());
}
```
