---
id: 68b7687dded630607aceccab
title: "チャレンジ 45: 完全平方数"
challengeType: 28
dashedName: challenge-45
---

# --description--

整数が与えられたとき、それが完全平方数かどうかを判定してください。

- 数が完全平方数であるとは、ある整数を自分自身で乗算してその数になる場合を指します。例えば、9は3を自分自身で乗算して得られるので完全平方数です。

# --hints--

`isPerfectSquare(9)` は `true` を返す必要があります。

```js
assert.isTrue(isPerfectSquare(9));
```

`isPerfectSquare(49)` は `true` を返す必要があります。

```js
assert.isTrue(isPerfectSquare(49));
```

`isPerfectSquare(1)` は `true` を返す必要があります。

```js
assert.isTrue(isPerfectSquare(1));
```

`isPerfectSquare(2)` は `false` を返す必要があります。

```js
assert.isFalse(isPerfectSquare(2));
```

`isPerfectSquare(99)` は `false` を返す必要があります。

```js
assert.isFalse(isPerfectSquare(99));
```

`isPerfectSquare(-9)` は `false` を返す必要があります。

```js
assert.isFalse(isPerfectSquare(-9));
```

`isPerfectSquare(0)` は `true` を返す必要があります。

```js
assert.isTrue(isPerfectSquare(0));
```

`isPerfectSquare(25281)` は `true` を返す必要があります。

```js
assert.isTrue(isPerfectSquare(25281));
```

# --seed--

## --seed-contents--

```js
function isPerfectSquare(n) {

  return n;
}
```

# --solutions--

```js
function isPerfectSquare(n) {
  if (n < 0) return false;
  const root = Math.floor(Math.sqrt(n));
  return root * root === n;
}
```
