---
id: 68b7687dded630607aceccb1
title: "チャレンジ 48: スパム検出器"
challengeType: 28
dashedName: challenge-48
---

# --description--

`"+A (BBB) CCC-DDDD"` の形式で与えられた電話番号について、各文字が次のように数字を表しています。

- `A` は国番号を表し、任意の桁数の数字が入ります。
- `BBB` はエリアコードを表し、常に3桁です。
- `CCC` と `DDDD` は市内番号を表し、それぞれ3桁と4桁です。

次の条件に基づいて、その番号がスパム番号かどうか判定してください。

- 国番号が2桁より長い、またはゼロ（`0`）で始まっていない。
- エリアコードが900より大きい、または200より小さい。
- 市内番号の最初の3桁の合計が、市内番号の最後の4桁の中に現れる。
- 同じ数字が4回以上連続している（フォーマット文字は無視）。

# --hints--

`isSpam("+0 (200) 234-0182")` は `false` を返すべきです。

```js
assert.isFalse(isSpam("+0 (200) 234-0182"));
```

`isSpam("+091 (555) 309-1922")` は `true` を返すべきです。

```js
assert.isTrue(isSpam("+091 (555) 309-1922"));
```

`isSpam("+1 (555) 435-4792")` は `true` を返すべきです。

```js
assert.isTrue(isSpam("+1 (555) 435-4792"));
```

`isSpam("+0 (955) 234-4364")` は `true` を返すべきです。

```js
assert.isTrue(isSpam("+0 (955) 234-4364"));
```

`isSpam("+0 (155) 131-6943")` は `true` を返すべきです。

```js
assert.isTrue(isSpam("+0 (155) 131-6943"));
```

`isSpam("+0 (555) 135-0192")` は `true` を返すべきです。

```js
assert.isTrue(isSpam("+0 (555) 135-0192"));
```

`isSpam("+0 (555) 564-1987")` は `true` を返すべきです。

```js
assert.isTrue(isSpam("+0 (555) 564-1987"));
```

`isSpam("+00 (555) 234-0182")` は `false` を返すべきです。

```js
assert.isFalse(isSpam("+00 (555) 234-0182"));
```

# --seed--

## --seed-contents--

```js
function isSpam(number) {

  return number;
}
```

# --solutions--

```js
function isSpam(number) {
  const match = number.match(/^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$/);
  const [, countryCode, areaCode, ccc, dddd] = match;
  const allDigits = countryCode + areaCode + ccc + dddd;
  
  if (countryCode.length > 2 || !countryCode.startsWith("0")) return true;

  const areaNum = parseInt(areaCode, 10);
  if (areaNum > 900 || areaNum < 200) return true;

  const sumCCC = ccc.split("").reduce((a, b) => a + parseInt(b), 0);
  if (dddd.includes(sumCCC.toString())) return true;

  if (/(\d)\1\1\1/.test(allDigits)) return true;

  return false;
}
```
