---
id: 68b7cadffed0e75a517da677
title: "チャレンジ 54: P@ssw0rd の強度！"
challengeType: 28
dashedName: challenge-54
---

# --description--

パスワードの文字列が与えられたら、その強度に応じて `"weak"`、`"medium"`、または `"strong"` を返してください。

パスワードは次のルールに基づいて評価されます：

- 8文字以上であること。
- 大文字と小文字の両方を含むこと。
- 少なくとも1つの数字を含むこと。
- 次のセットの中から少なくとも1つの特殊文字を含むこと：`!`、`@`、`#`、`$`、`%`、`^`、`&`、または `*`。

ルールを2つ未満しか満たさない場合は `"weak"` を返してください。
ルールを2つまたは3つ満たす場合は `"medium"` を返してください。
ルールをすべて4つ満たす場合は `"strong"` を返してください。

# --hints--

`checkStrength("123456")` は `"weak"` を返すべきです。

```js
assert.equal(checkStrength("123456"), "weak");
```

`checkStrength("pass!!!")` は `"weak"` を返すべきです。

```js
assert.equal(checkStrength("pass!!!"), "weak");
```

`checkStrength("Qwerty")` は `"weak"` を返すべきです。

```js
assert.equal(checkStrength("Qwerty"), "weak");
```

`checkStrength("PASSWORD")` は `"weak"` を返すべきです。

```js
assert.equal(checkStrength("PASSWORD"), "weak");
```

`checkStrength("PASSWORD!")` は `"medium"` を返すべきです。

```js
assert.equal(checkStrength("PASSWORD!"), "medium");
```

`checkStrength("PassWord%^!")` は `"medium"` を返すべきです。

```js
assert.equal(checkStrength("PassWord%^!"), "medium");
```

`checkStrength("qwerty12345")` は `"medium"` を返すべきです。

```js
assert.equal(checkStrength("qwerty12345"), "medium");
```

`checkStrength("S3cur3P@ssw0rd")` は `"strong"` を返すべきです。

```js
assert.equal(checkStrength("S3cur3P@ssw0rd"), "strong");
```

`checkStrength("C0d3&Fun!")` は `"strong"` を返すべきです。

```js
assert.equal(checkStrength("C0d3&Fun!"), "strong");
```

# --seed--

## --seed-contents--

```js
function checkStrength(password) {

  return password;
}
```

# --solutions--

```js
function checkStrength(password) {
  let rulesMet = 0;

  if (password.length >= 8) rulesMet++;
  if (/[a-z]/.test(password) && /[A-Z]/.test(password)) rulesMet++;
  if (/\d/.test(password)) rulesMet++;
  if (/[!@#$%^&*]/.test(password)) rulesMet++;

  if (rulesMet < 2) return "weak";
  if (rulesMet < 4) return "medium";
  return "strong";
}
```
