---
id: 68b7cadffed0e75a517da675
title: "チャレンジ 53: 10進数から2進数へ"
challengeType: 28
dashedName: challenge-53
---

# --description--

0以上の整数が与えられたら、その2進数表現を文字列として返してください。

2進数は数字の`0`と`1`だけを使ってどんな数でも表します。10進数を2進数に変換するには、その数を`2`で繰り返し割り、余りを記録します。数が0になるまで繰り返してください。最後に記録した余りから最初の余りまでを読みます。例えば、`12`を2進数に変換すると：

```mathml
12 ÷ 2 = 6 remainder 0
6 ÷ 2 = 3 remainder 0
3 ÷ 2 = 1 remainder 1
1 ÷ 2 = 0 remainder 1
```

`12`の2進数表現は`1100`です。

# --hints--

`toBinary(5)`は`"101"`を返すべきです。

```js
assert.equal(toBinary(5), "101");
```

`toBinary(12)`は`"1100"`を返すべきです。

```js
assert.equal(toBinary(12), "1100");
```

`toBinary(50)`は`"110010"`を返すべきです。

```js
assert.equal(toBinary(50), "110010");
```

`toBinary(99)`は`"1100011"`を返すべきです。

```js
assert.equal(toBinary(99), "1100011");
```

# --seed--

## --seed-contents--

```js
function toBinary(decimal) {

  return decimal;
}
```

# --solutions--

```js
function toBinary(decimal) {
  if (decimal === 0) return "0";
  let binary = "";
  while (decimal > 0) {
    binary = (decimal % 2) + binary;
    decimal = Math.floor(decimal / 2);
  }
  return binary;
}
```
