---
id: 68b7687dded630607aceccad
title: "チャレンジ 46: 2番目に大きい数"
challengeType: 28
dashedName: challenge-46
---

# --description--

配列が与えられたとき、2番目に大きい異なる数を返してください。

# --hints--

`secondLargest([1, 2, 3, 4])` は `3` を返すべきです。

```js
assert.equal(secondLargest([1, 2, 3, 4]), 3);
```

`secondLargest([20, 139, 94, 67, 31])` は `94` を返すべきです。

```js
assert.equal(secondLargest([20, 139, 94, 67, 31]), 94);
```

`secondLargest([2, 3, 4, 6, 6])` は `4` を返すべきです。

```js
assert.equal(secondLargest([2, 3, 4, 6, 6]), 4);
```

`secondLargest([10, -17, 55.5, 44, 91, 0])` は `55.5` を返すべきです。

```js
assert.equal(secondLargest([10, -17, 55.5, 44, 91, 0]), 55.5);
```

`secondLargest([1, 0, -1, 0, 1, 0, -1, 1, 0])` は `0` を返すべきです。

```js
assert.equal(secondLargest([1, 0, -1, 0, 1, 0, -1, 1, 0]), 0);
```

# --seed--

## --seed-contents--

```js
function secondLargest(arr) {

  return arr;
}
```

# --solutions--

```js
function secondLargest(arr) {
  const unique = [...new Set(arr)];
  unique.sort((a, b) => b - a);
  return unique[1];
}
```
