---
id: 68b7687dded630607aceccaf
title: "チャレンジ 47: スピード違反の検出"
challengeType: 28
dashedName: challenge-47
---

# --description--

車両の走行速度を表す数値の配列と制限速度を表す数値が与えられたとき、スピード違反をしていた車両の数と、その車両たちの制限速度超過の平均値を含む2つの要素からなる配列を返してください。

- スピード違反をしている車両がいなかった場合は `[0, 0]` を返してください。

# --hints--

`speeding([50, 60, 55], 60)` は `[0, 0]` を返すべきです。

```js
assert.deepEqual(speeding([50, 60, 55], 60), [0, 0]);
```

`speeding([58, 50, 60, 55], 55)` は `[2, 4]` を返すべきです。

```js
assert.deepEqual(speeding([58, 50, 60, 55], 55), [2, 4]);
```

`speeding([61, 81, 74, 88, 65, 71, 68], 70)` は `[4, 8.5]` を返すべきです。

```js
assert.deepEqual(speeding([61, 81, 74, 88, 65, 71, 68], 70), [4, 8.5]);
```

`speeding([100, 105, 95, 102], 100)` は `[2, 3.5]` を返すべきです。

```js
assert.deepEqual(speeding([100, 105, 95, 102], 100), [2, 3.5]);
```

`speeding([40, 45, 44, 50, 112, 39], 55)` は `[1, 57]` を返すべきです。

```js
assert.deepEqual(speeding([40, 45, 44, 50, 112, 39], 55), [1, 57]);
```

# --seed--

## --seed-contents--

```js
function speeding(speeds, limit) {

  return speeds;
}
```

# --solutions--

```js
function speeding(speeds, limit) {
  let speeding = 0;
  let totalOver = 0;

  for (let speed of speeds) {
    if (speed > limit) {
      speeding++;
      totalOver += (speed - limit);
    }
  }

  if (speeding === 0) return [0, 0];

  return [speeding, totalOver / speeding];
}
```
