---
id: 68b7cadffed0e75a517da671
title: "チャレンジ 51: 電話番号フォーマッター"
challengeType: 28
dashedName: challenge-51
---

# --description--

11桁の数字の文字列が与えられたら、その文字列をこの形式の電話番号として返してください：`"+D (DDD) DDD-DDDD"`。

# --hints--

`formatNumber("05552340182")` は `"+0 (555) 234-0182"` を返すべきです。

```js
assert.equal(formatNumber("05552340182"), "+0 (555) 234-0182");
```

`formatNumber("15554354792")` は `"+1 (555) 435-4792"` を返すべきです。

```js
assert.equal(formatNumber("15554354792"), "+1 (555) 435-4792");
```

# --seed--

## --seed-contents--

```js
function formatNumber(number) {

  return number;
}
```

# --solutions--

```js
function formatNumber(number) {
  const country = number[0];
  const area = number.slice(1, 4);
  const prefix = number.slice(4, 7);
  const line = number.slice(7);

  return `+${country} (${area}) ${prefix}-${line}`;
}
```
