---
id: 68b7cadffed0e75a517da673
title: "チャレンジ 52: 2進数から10進数へ"
challengeType: 28
dashedName: challenge-52
---

# --description--

2進数を表す文字列が与えられたら、その10進数の値を数値として返してください。

2進数は数字の`0`と`1`だけを使ってどんな数でも表します。2進数を10進数に変換するには、各桁を`2`のべき乗で乗算して、それらをすべて足し合わせます。右端の桁を`2^0`で乗算し、その左の桁を`2^1`で乗算し、以下同様に続けます。すべての桁を`2`のべき乗で乗算したら、その結果を合計してください。

たとえば、2進数の`101`は10進数で`5`に等しいです。なぜなら：

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`toDecimal("101")`は`5`を返すべきです。

```js
assert.equal(toDecimal("101"), 5);
```

`toDecimal("1010")`は`10`を返すべきです。

```js
assert.equal(toDecimal("1010"), 10);
```

`toDecimal("10010")`は`18`を返すべきです。

```js
assert.equal(toDecimal("10010"), 18);
```

`toDecimal("1010101")`は`85`を返すべきです。

```js
assert.equal(toDecimal("1010101"), 85);
```

# --seed--

## --seed-contents--

```js
function toDecimal(binary) {

  return binary;
}
```

# --solutions--

```js
function toDecimal(binary) {
  let decimal = 0;
  for (let i = 0; i < binary.length; i++) {
    const bit = parseInt(binary[binary.length - 1 - i], 10);
    decimal += bit * Math.pow(2, i);
  }
  return decimal;
}
```
