---
id: 68c1a929005bf54d342aa8d4
title: "チャレンジ 57: スペースウィーク3日目: 電話をかける"
challengeType: 28
dashedName: challenge-57
---

# --description--

スペースウィークの3日目では、通信ルート上の自分、衛星、そして母星の間の距離（キロメートル単位）を表す配列が与えられます。次の条件を使って、メッセージがルートを通って目的の惑星に届くまでにかかる時間を求めてください。

- 配列の最初の値は、自分の位置から最初の衛星までの距離です。
- 最後の値を除く各値は、次の衛星までの距離です。
- 配列の最後の値は、前の衛星から母星までの距離です。
- メッセージは秒速300,000キロメートルで移動します。
- メッセージが通過する各衛星は0.5秒の伝送遅延を追加します。
- 小数点以下4桁に丸め、末尾のゼロは削除して数値を返してください。

# --hints--

`sendMessage([300000, 300000])` は `2.5` を返すべきです。

```js
assert.equal(sendMessage([300000, 300000]), 2.5);
```

`sendMessage([384400, 384400])` は `3.0627` を返すべきです。

```js
assert.equal(sendMessage([384400, 384400]), 3.0627);
```

`sendMessage([54600000, 54600000])` は `364.5` を返すべきです。

```js
assert.equal(sendMessage([54600000, 54600000]), 364.5);
```

`sendMessage([1000000, 500000000, 1000000])` は `1674.3333` を返すべきです。

```js
assert.equal(sendMessage([1000000, 500000000, 1000000]), 1674.3333);
```

`sendMessage([10000, 21339, 50000, 31243, 10000])` は `2.4086` を返すべきです。

```js
assert.equal(sendMessage([10000, 21339, 50000, 31243, 10000]), 2.4086);
```

`sendMessage([802101, 725994, 112808, 3625770, 481239])` は `21.1597` を返すべきです。

```js
assert.equal(sendMessage([802101, 725994, 112808, 3625770, 481239]), 21.1597);
```

# --seed--

## --seed-contents--

```js
function sendMessage(route) {

  return route;
}
```

# --solutions--

```js
function sendMessage(route) {
  let totalDistance = route.reduce((a, d) => a += d, 0);

  const delay = (route.length - 1) * 0.5
  const time = totalDistance / 300000
  const total = time + delay;
  return Number(total.toFixed(4))
}
```
