---
id: 68c1a929005bf54d342aa8d2
title: "チャレンジ 55: 宇宙週間 第1日目: 恒星の分類"
challengeType: 28
dashedName: challenge-55
---

# --description--

10月4日は世界宇宙週間の始まりです。これからの7日間は天文学をテーマにしたコーディングチャレンジが続きます。

今日のチャレンジでは、星の表面温度がケルビン（K）で与えられ、その温度範囲に基づいて恒星の分類を判定してください。

- `"O"`: 30,000 K以上
- `"B"`: 10,000 K〜29,999 K
- `"A"`: 7,500 K〜9,999 K
- `"F"`: 6,000 K〜7,499 K
- `"G"`: 5,200 K〜5,999 K
- `"K"`: 3,700 K〜5,199 K
- `"M"`: 0 K〜3,699 K

- 与えられた星の分類を返してください。

# --hints--

`classification(5778)` は `"G"` を返すべきです。

```js
assert.equal(classification(5778), "G");
```

`classification(2400)` は `"M"` を返すべきです。

```js
assert.equal(classification(2400), "M");
```

`classification(9999)` は `"A"` を返すべきです。

```js
assert.equal(classification(9999), "A");
```

`classification(3700)` は `"K"` を返すべきです。

```js
assert.equal(classification(3700), "K");
```

`classification(3699)` は `"M"` を返すべきです。

```js
assert.equal(classification(3699), "M");
```

`classification(210000)` は `"O"` を返すべきです。

```js
assert.equal(classification(210000), "O");
```

`classification(6000)` は `"F"` を返すべきです。

```js
assert.equal(classification(6000), "F");
```

`classification(11432)` は `"B"` を返すべきです。

```js
assert.equal(classification(11432), "B");
```

# --seed--

## --seed-contents--

```js
function classification(temp) {

  return temp;
}
```

# --solutions--

```js
function classification(temp) {
  if (temp >= 30000) return "O";
  if (temp >= 10000) return "B";
  if (temp >= 7500) return "A";
  if (temp >= 6000) return "F";
  if (temp >= 5200) return "G";
  if (temp >= 3700) return "K";
  return "M";
}
```
