---
id: 68c1a929005bf54d342aa8d3
title: "チャレンジ 56: スペースウィーク 2日目: 太陽系外惑星の探索"
challengeType: 28
dashedName: challenge-56
---

# --description--

スペースウィークの2日目では、各文字が星の光度の読み取り値を表す文字列が与えられます。トランジット法を使って太陽系外惑星が検出されたかどうかを判定してください。トランジット法とは、惑星が星の前を通過して観測される光度が減少する現象のことです。

- 光度の読み取り値は `0-9` と `A-Z` の文字のみで構成され、それぞれの読み取り値は次の数値に対応します：
- 文字 `0-9` は光度レベル `0-9` に対応します。
- 文字 `A-Z` は光度レベル `10-35` に対応します。

星に太陽系外惑星があるとみなす条件は、任意の単一の読み取り値が全読み取り値の平均の80％以下であることです。例えば、星の平均光度が10の場合、任意の読み取り値が8以下であれば太陽系外惑星があるとみなされます。

# --hints--

`hasExoplanet("665544554")` は `false` を返すべきです。

```js
assert.isFalse(hasExoplanet("665544554"));
```

`hasExoplanet("FGFFCFFGG")` は `true` を返すべきです。

```js
assert.isTrue(hasExoplanet("FGFFCFFGG"));
```

`hasExoplanet("MONOPLONOMONPLNOMPNOMP")` は `false` を返すべきです。

```js
assert.isFalse(hasExoplanet("MONOPLONOMONPLNOMPNOMP"));
```

`hasExoplanet("FREECODECAMP")` は `true` を返すべきです。

```js
assert.isTrue(hasExoplanet("FREECODECAMP"));
```

`hasExoplanet("9AB98AB9BC98A")` は `false` を返すべきです。

```js
assert.isFalse(hasExoplanet("9AB98AB9BC98A"));
```

`hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` は `true` を返すべきです。

```js
assert.isTrue(hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"));
```

# --seed--

## --seed-contents--

```js
function hasExoplanet(readings) {

  return readings;
}
```

# --solutions--

```js
function hasExoplanet(readings) {
  let total = 0;
  const values = readings.split('').map(c => parseInt(c, 36));
  
  total = values.reduce((sum, v) => sum + v, 0);
  const average = total / values.length;
  const threshold = average * 0.8;

  return values.some(v => v <= threshold);
}
```
