---
id: 68c1a929005bf54d342aa8d5
title: "チャレンジ 58: スペースウィーク第4日目: 着陸地点"
challengeType: 28
dashedName: challenge-58
---

# --description--

スペースウィークの4日目では、ローバーの潜在的な着陸地点を表す数値のマトリックス（配列の配列）が与えられます。次のルールに基づいて最も安全な着陸地点を見つけてください。

- マトリックスの各地点には`0-9`からの数値が含まれます（`0-9`を含む）。
- `0`は潜在的な着陸地点を表します。
- `0`以外の数値は着陸には危険すぎます。数値が大きいほど危険度が高くなります。
- 最も安全な地点は、周囲のセル（最大4つの隣接セル、斜めは無視）の合計危険度が最も低い`0`のセルと定義されます。
- 範囲外の隣接セルは無視してください（角や端は隣接セルが少なくなります）。
- 最も安全な着陸地点のインデックスを返してください。最も安全な地点は必ず1つだけ存在します。

例えば、次のような場合：

```js
[
  [1, 0],
  [2, 0]
]
```

最初の配列の`0`のインデックスである`[0, 1]`を返してください。

# --hints--

`findLandingSpot([[1, 0], [2, 0]])`は`[0, 1]`を返すべきです。

```js
assert.deepEqual(findLandingSpot([[1, 0], [2, 0]]), [0, 1]);
```

`findLandingSpot([[9, 0, 3], [7, 0, 4], [8, 0, 5]])`は`[1, 1]`を返すべきです。

```js
assert.deepEqual(findLandingSpot([[9, 0, 3], [7, 0, 4], [8, 0, 5]]), [1, 1]);
```

`findLandingSpot([[1, 2, 1], [0, 0, 2], [3, 0, 0]])`は`[2, 2]`を返すべきです。

```js
assert.deepEqual(findLandingSpot([[1, 2, 1], [0, 0, 2], [3, 0, 0]]), [2, 2]);
```

`findLandingSpot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]])`は`[2, 1]`を返すべきです。

```js
assert.deepEqual(findLandingSpot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]]), [2, 1]);
```

# --seed--

## --seed-contents--

```js
function findLandingSpot(matrix) {

  return matrix;
}
```

# --solutions--

```js
function findLandingSpot(matrix) {
  let bestSpot = null;
  let lowestNeighborSum = Infinity;

  for (let i = 0; i < matrix.length; i++) {
    for (let j = 0; j < matrix[i].length; j++) {

      if (matrix[i][j] === 0) {
        let currentNeighborSum = 0;

        if (i > 0) currentNeighborSum += matrix[i - 1][j];
        if (j < matrix[i].length - 1) currentNeighborSum += matrix[i][j + 1];
        if (i < matrix.length - 1) currentNeighborSum += matrix[i + 1][j];
        if (j > 0) currentNeighborSum += matrix[i][j - 1];

        if (currentNeighborSum < lowestNeighborSum) {
          lowestNeighborSum = currentNeighborSum;
          bestSpot = [i, j];
        }
      }
    }
  }

  return bestSpot;
}
```
