---
id: 68c1a929005bf54d342aa8d6
title: "チャレンジ 59: 宇宙週間 第5日目: ゴールディロックスゾーン"
challengeType: 28
dashedName: challenge-59
---

# --description--

宇宙週間の5日目では、恒星の「ゴールディロックスゾーン」を計算します。これは、液体の水が存在するのに「ちょうど良い」条件が整っている恒星の周囲の領域です。

恒星の質量が与えられたら、そのゴールディロックスゾーンの開始距離と終了距離を天文単位で示す配列を返してください。

ゴールディロックスゾーンを計算するには：

1. 恒星の光度を、質量を3.5乗した値で求めます。
2. ゾーンの開始距離は、その光度の平方根に0.95を掛けた値です。
3. ゾーンの終了距離は、その光度の平方根に1.37を掛けた値です。

- 距離は小数点以下2桁に丸めて返してください。

例えば、質量が`1`の場合、`[0.95, 1.37]`を返します。

# --hints--

`goldilocksZone(1)`は`[0.95, 1.37]`を返すべきです。

```js
assert.deepEqual(goldilocksZone(1), [0.95, 1.37]);
```

`goldilocksZone(0.5)`は`[0.28, 0.41]`を返すべきです。

```js
assert.deepEqual(goldilocksZone(0.5), [0.28, 0.41]);
```

`goldilocksZone(6)`は`[21.85, 31.51]`を返すべきです。

```js
assert.deepEqual(goldilocksZone(6), [21.85, 31.51]);
```

`goldilocksZone(3.7)`は`[9.38, 13.52]`を返すべきです。

```js
assert.deepEqual(goldilocksZone(3.7), [9.38, 13.52]);
```

`goldilocksZone(20)`は`[179.69, 259.13]`を返すべきです。

```js
assert.deepEqual(goldilocksZone(20), [179.69, 259.13]);
```

# --seed--

## --seed-contents--

```js
function goldilocksZone(mass) {

  return mass;
}
```

# --solutions--

```js
function goldilocksZone(mass) {
  const luminosity = Math.pow(mass, 3.5);
  const start = 0.95 * Math.sqrt(luminosity);
  const end = 1.37 * Math.sqrt(luminosity);
  return [Number(start.toFixed(2)), Number(end.toFixed(2))];
}
```
