---
id: 68c497f3aaefc9fd9f1b0e26
title: "チャレンジ 62: 16進数から10進数へ"
challengeType: 28
dashedName: challenge-62
---

# --description--

16進数（基数16）を表す文字列が与えられたら、その10進数（基数10）の値を整数として返してください。

16進数は16個の数字を使う数値システムです：

- `0-9` は `0` から `9` までの値を表します。
- `A-F` は `10` から `15` までの値を表します。

部分的な変換テーブルは以下の通りです：

| 16進数 | 10進数 |
|:-------:|:-------:|
| 0       | 0       |
| 1       | 1       |
| ...     | ...     |
| 9       | 9       |
| A       | 10      |
| ...     | ...     |
| F       | 15      |
| 10      | 16      |
| ...     | ...     |
| 9F      | 159     |
| A0      | 160     |
| ...     | ...     |
| FF      | 255     |
| 100     | 256     |

- 文字列には `0–9` と `A–F` の文字のみが含まれます。

# --hints--

`hexToDecimal("A")` は `10` を返すべきです。

```js
assert.equal(hexToDecimal("A"), 10);
```

`hexToDecimal("15")` は `21` を返すべきです。

```js
assert.equal(hexToDecimal("15"), 21);
```

`hexToDecimal("2E")` は `46` を返すべきです。

```js
assert.equal(hexToDecimal("2E"), 46);
```

`hexToDecimal("FF")` は `255` を返すべきです。

```js
assert.equal(hexToDecimal("FF"), 255);
```

`hexToDecimal("A3F")` は `2623` を返すべきです。

```js
assert.equal(hexToDecimal("A3F"), 2623);
```

# --seed--

## --seed-contents--

```js
function hexToDecimal(hex) {

  return hex;
}
```

# --solutions--

```js
function hexToDecimal(hex) {

  return parseInt(hex, 16);
}
```
