---
id: 68c497f3aaefc9fd9f1b0e25
title: "チャレンジ 61: スペースウィーク第7日目: 発射燃料"
challengeType: 28
dashedName: challenge-61
---

# --description--

スペースウィークの最終日には、軌道に送るペイロードの質量（キログラム単位）が与えられます。以下のルールに従って、ペイロードを軌道に送るために必要な燃料の量を求めてください。

- ロケットは持ち上げる質量5 kgごとに1 kgの燃料が必要です。
- 燃料自体にも質量があります。燃料を追加すると持ち上げる質量が増え、それによりさらに燃料が必要になり、質量が増え…という繰り返しになります。
- 必要な燃料の総量を計算するには、まずペイロードの質量から燃料の必要量を計算し、その燃料を総質量に加え、再度計算します。このプロセスを追加で必要な燃料が1 kg未満になるまで繰り返してください。
- ロケット自体の質量は無視します。ペイロードとその燃料を持ち上げるために必要な燃料のみを計算してください。

例えば、ペイロードの質量が50 kgの場合、持ち上げるために10 kgの燃料が必要です（ペイロード質量 ÷ 5）。これにより総質量は60 kgとなり、さらに12 kgの燃料が必要になります（追加で2 kg）。総質量は62 kgとなり、12.4 kgの燃料が必要になります（追加で0.4 kg）。追加の燃料が1 kg未満なのでここで計算を止めます。持ち上げる総質量は62.4 kgで、そのうち50 kgが初期のペイロード、12.4 kgが燃料です。

- 必要な燃料の量は小数点以下1桁に丸めて返してください。

# --hints--

`launchFuel(50)` should return `12.4`.

```js
assert.equal(launchFuel(50), 12.4);
```

`launchFuel(500)` should return `124.8`.

```js
assert.equal(launchFuel(500), 124.8);
```

`launchFuel(243)` should return `60.7`.

```js
assert.equal(launchFuel(243), 60.7);
```

`launchFuel(11000)` should return `2749.8`.

```js
assert.equal(launchFuel(11000), 2749.8);
```

`launchFuel(6214)` should return `1553.4`.

```js
assert.equal(launchFuel(6214), 1553.4);
```

# --seed--

## --seed-contents--

```js
function launchFuel(payload) {

  return payload;
}
```

# --solutions--

```js
function launchFuel(payload) {
  let totalMass = payload;
  let additionalFuel = totalMass / 5;
  let totalFuel = additionalFuel;

  while (additionalFuel >= 1) {
    totalMass += additionalFuel;
    additionalFuel = (totalMass / 5) - totalFuel;
    totalFuel += additionalFuel;
  }

  return Number(totalFuel.toFixed(1));
}
```
