---
id: 68c497f3aaefc9fd9f1b0e24
title: "チャレンジ 60: 宇宙週間6日目: 月の満ち欠け"
challengeType: 28
dashedName: challenge-60
---

# --description--

宇宙週間の6日目では、`"YYYY-MM-DD"` の形式で日付が与えられ、その日の月の満ち欠けを次のルールに従って判定してください。

28日間の簡略化した月の周期を使い、4つの等しいフェーズに分けます。

- `"New"`: 1日目から7日目まで
- `"Waxing"`: 8日目から14日目まで
- `"Full"`: 15日目から21日目まで
- `"Waning"`: 22日目から28日目まで

28日目の後は周期が繰り返され、1日目は新月となります。

- `"2000-01-06"` を基準の新月（周期の1日目）として使い、与えられた日のフェーズを判定してください。
- 基準日より前の日付は与えられません。
- 正しいフェーズを文字列で返してください。

**注意:** 1日目は新月の日を表し、前回の新月から0日経過していることを意味します。

# --hints--

`moonPhase("2000-01-12")` は `"New"` を返すべきです。

```js
assert.equal(moonPhase("2000-01-12"), "New");
```

`moonPhase("2000-01-13")` は `"Waxing"` を返すべきです。

```js
assert.equal(moonPhase("2000-01-13"), "Waxing");
```

`moonPhase("2014-10-15")` は `"Full"` を返すべきです。

```js
assert.equal(moonPhase("2014-10-15"), "Full");
```

`moonPhase("2012-10-21")` は `"Waning"` を返すべきです。

```js
assert.equal(moonPhase("2012-10-21"), "Waning");
```

`moonPhase("2022-12-14")` は `"New"` を返すべきです。

```js
assert.equal(moonPhase("2022-12-14"), "New");
```

# --seed--

## --seed-contents--

```js
function moonPhase(dateString) {

  return dateString;
}
```

# --solutions--

```js
function moonPhase(dateString) {
  const ONE_DAY_IN_MS = 1000 * 60 * 60 * 24;
  const refDate = new Date("2000-01-06");
  const targetDate = new Date(dateString);
  const diffMs = targetDate - refDate;
  const diffDays = diffMs / ONE_DAY_IN_MS;
  const phaseDay = (diffDays % 28 + 1);

  if (phaseDay <= 7) {
    return "New";
  } else if (phaseDay <= 14) {
    return "Waxing";
  } else if (phaseDay <= 21) {
    return "Full"
  } else {
    return "Waning";
  }
}

```
