---
id: 68cae5b538ff798bbd4da009
title: "チャレンジ 71: チップ計算機"
challengeType: 28
dashedName: challenge-71
---

# --description--

食事の価格とカスタムのチップ率が与えられたとき、15%、20%、およびカスタムの3つのチップ額を含む配列を返してください。

- 価格は `"$N.NN"` の形式で与えられます。
- カスタムのチップ率は `"25%"` の形式で与えられます。
- 金額は同じ `"$N.NN"` 形式で、小数点以下2桁に丸めて返してください。

たとえば、`"$10.00"` の食事価格と `"25%"` のカスタムチップ値が与えられた場合、`["$1.50", "$2.00", "$2.50"]` を返します。

# --hints--

`calculateTips("$10.00", "25%")` は `["$1.50", "$2.00", "$2.50"]` を返すべきです。

```js
assert.deepEqual(calculateTips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"]);
```

`calculateTips("$89.67", "26%")` は `["$13.45", "$17.93", "$23.31"]` を返すべきです。

```js
assert.deepEqual(calculateTips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"]);
```

`calculateTips("$19.85", "9%")` は `["$2.98", "$3.97", "$1.79"]` を返すべきです。

```js
assert.deepEqual(calculateTips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"]);
```

# --seed--

## --seed-contents--

```js
function calculateTips(mealPrice, customTip) {

  return mealPrice;
}
```

# --solutions--

```js
function calculateTips(mealPrice, customTip) {
  const meal = parseFloat(mealPrice.slice(1));
  const customPercent = parseFloat(customTip.slice(0, -1));
  const tips = [15, 20, customPercent];

  return tips.map(percent => {
    const amount = (meal * percent / 100).toFixed(2);
    return `$${amount}`;
  });
}
```
