---
id: 68cae5b538ff798bbd4da003
title: "チャレンジ 65: 文字列の出現回数"
challengeType: 28
dashedName: challenge-65
---

# --description--

2つの文字列が与えられたとき、2番目の文字列が1番目の文字列に何回現れるかを判定してください。

- パターン文字列は1番目の文字列の中で重なって現れる場合があります。例えば、`"aaa"`には`"aa"`が2回含まれています。最初の2つの`a`と2番目の2つの`a`です。

# --hints--

`count('abcdefg', 'def')`は`1`を返すべきです。

```js
assert.equal(count('abcdefg', 'def'), 1);
```

`count('hello', 'world')`は`0`を返すべきです。

```js
assert.equal(count('hello', 'world'), 0);
```

`count('mississippi', 'iss')`は`2`を返すべきです。

```js
assert.equal(count('mississippi', 'iss'), 2);
```

`count('she sells seashells by the seashore', 'sh')`は`3`を返すべきです。

```js
assert.equal(count('she sells seashells by the seashore', 'sh'), 3);
```

`count('101010101010101010101', '101')`は`10`を返すべきです。

```js
assert.equal(count('101010101010101010101', '101'), 10);
```

# --seed--

## --seed-contents--

```js
function count(text, pattern) {

  return text;
}
```

# --solutions--

```js
function count(text, pattern) {
  let occurrences = 0;

  for (let i = 0; i <= text.length - pattern.length; i++) {
    if (text.slice(i, i + pattern.length) === pattern) {
      occurrences++;
    }
  }

  return occurrences;
}
```
