---
id: 68cae5b538ff798bbd4da006
title: "チャレンジ 68: クレジットカードのマスキング"
challengeType: 28
dashedName: challenge-68
---

# --description--

クレジットカード番号の文字列が与えられたとき、以下の条件に従ってマスクしたバージョンを返してください。

- 文字列は4桁ずつの数字が4セット含まれており（`0-9`）、すべてのセットは単一のスペースまたは単一のハイフン（`-`）で区切られています。
- 最後の4桁を除くすべての数字をアスタリスク（`*`）に置き換えます。
- 残りの文字は変更しません。

例えば、`"4012-8888-8888-1881"` が与えられた場合、`"****-****-****-1881"` を返してください。

# --hints--

`mask("4012-8888-8888-1881")` は `"****-****-****-1881"` を返すべきです。

```js
assert.equal(mask("4012-8888-8888-1881"), "****-****-****-1881");
```

`mask("5105 1051 0510 5100")` は `"**** **** **** 5100"` を返すべきです。

```js
assert.equal(mask("5105 1051 0510 5100"), "**** **** **** 5100");
```

`mask("6011 1111 1111 1117")` は `"**** **** **** 1117"` を返すべきです。

```js
assert.equal(mask("6011 1111 1111 1117"), "**** **** **** 1117");
```

`mask("2223-0000-4845-0010")` は `"****-****-****-0010"` を返すべきです。

```js
assert.equal(mask("2223-0000-4845-0010"), "****-****-****-0010");
```

# --seed--

## --seed-contents--

```js
function mask(card) {

  return card;
}
```

# --solutions--

```js
function mask(card) {
  const split = card.split(card.includes('-') ? '-' : ' ');
  
  return card.includes('-') ?
    `****-****-****-${split[3]}` :
    `**** **** **** ${split[3]}`;
}
```
