---
id: 68cae5b538ff798bbd4da004
title: "チャレンジ 66: HTML タグストリッパー"
challengeType: 28
dashedName: challenge-66
---

# --description--

HTMLコードの文字列が与えられたら、タグを取り除いてプレーンテキストの内容を返してください。

- 入力の文字列は有効なHTMLのみを含みます。
- HTMLタグはネストされている場合があります。
- タグとその属性はすべて取り除いてください。

例えば、`'<a href="#">Click here</a>'` は `"Click here"` を返すべきです。

# --hints--

`stripTags('<a href="#">Click here</a>')` は `"Click here"` を返すべきです。

```js
assert.equal(stripTags('<a href="#">Click here</a>'), "Click here");
```

`stripTags('<p class="center">Hello <b>World</b>!</p>')` は `"Hello World!"` を返すべきです。

```js
assert.equal(stripTags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!");
```

`stripTags('<img src="cat.jpg" alt="Cat">')` は空の文字列（`""`）を返すべきです。

```js
assert.equal(stripTags('<img src="cat.jpg" alt="Cat">'), "");
```

`stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` は `sectionsection` を返すべきです。

```js
assert.equal(stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection");
```

# --seed--

## --seed-contents--

```js
function stripTags(html) {

  return html;
}
```

# --solutions--

```js
function stripTags(html) {
  return html.replace(/<[^>]*>/g, '');
}
```
