---
id: 68cae5b538ff798bbd4da001
title: "チャレンジ 63: 言葉の戦い"
challengeType: 28
dashedName: challenge-63
---

# --description--

あなたのチームと対戦チームを表す2つの文が与えられます。各単語が対応する相手チームの単語と戦うとき、次のルールに従ってどちらのチームが勝つか判定してください。

- 与えられる文は常に同じ単語数です。
- 単語は1つのスペースで区切られ、文字のみで構成されます。
- 各単語の値は、その単語の文字の合計値です。
- 文字`a`から`z`は値`1`から`26`に対応します。例えば、`a`は`1`であり、`z`は`26`です。
- 大文字は文字の値を2倍にします。例えば、`A`は`2`であり、`Z`は`52`です。
- 単語は順番に戦います：あなたのチームの最初の単語が対戦チームの最初の単語と戦い、以下同様です。
- 単語の値が相手の単語の値より大きい場合、その単語が勝ちます。
- 勝った単語の数が多いチームが勝者です。

あなたのチームが勝者の場合は`"We win"`を、負けた場合は`"We lose"`を、両チームの勝ち数が同じ場合は`"Draw"`を返してください。

# --hints--

`battle("hello world", "hello word")`は`"We win"`を返すべきです。

```js
assert.equal(battle("hello world", "hello word"), "We win");
```

`battle("Hello world", "hello world")`は`"We win"`を返すべきです。

```js
assert.equal(battle("Hello world", "hello world"), "We win");
```

`battle("lorem ipsum", "kitty ipsum")`は`"We lose"`を返すべきです。

```js
assert.equal(battle("lorem ipsum", "kitty ipsum"), "We lose");
```

`battle("hello world", "world hello")`は`"Draw"`を返すべきです。

```js
assert.equal(battle("hello world", "world hello"), "Draw");
```

`battle("git checkout", "git switch")`は`"We win"`を返すべきです。

```js
assert.equal(battle("git checkout", "git switch"), "We win");
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")`は`"We lose"`を返すべきです。

```js
assert.equal(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose");
```

`battle("We must never surrender", "Our team must win")`は`"Draw"`を返すべきです。

```js
assert.equal(battle("We must never surrender", "Our team must win"), "Draw");
```

# --seed--

## --seed-contents--

```js
function battle(ourTeam, opponent) {

  return ourTeam;
}
```

# --solutions--

```js
function getCharacterValue(char) {
  const lowerCaseValue = char.toLowerCase().charCodeAt(0) - 96;
  return char >= 'A' && char <= 'Z' ? lowerCaseValue * 2 : lowerCaseValue;
}

function getWordValue(word) {
  let wordValue = 0;
  for (let i=0; i<word.length; i++) {
    wordValue += getCharacterValue(word[i])
  }

  return wordValue;
}

function battle(ourTeam, opponent) {
  const myWords = ourTeam.split(' ');
  const opponentWords = opponent.split(' ');

  let myWins = 0, opponentWins = 0;

  for (let i=0; i<myWords.length; i++) {
    const myWordValue = getWordValue(myWords[i]);
    const opponentWordValue = getWordValue(opponentWords[i])

    if (myWordValue > opponentWordValue) myWins++;
    if (opponentWordValue > myWordValue) opponentWins++;
  }

  if (myWins > opponentWins) return 'We win';
  if (opponentWins > myWins) return 'We lose';
  return 'Draw';
}
```
