---
id: 68cae5b538ff798bbd4da002
title: "チャレンジ 64: 24時間表記から12時間表記へ"
challengeType: 28
dashedName: challenge-64
---

# --description--

24時間表記の`"HHMM"`で表された時刻の文字列が与えられたら、その時刻を同等の12時間表記の`"H:MM AM"`または`"H:MM PM"`で返してください。

- 与えられる入力は常に24時間表記の4桁の文字列で、`"0000"`から`"2359"`の範囲です。

# --hints--

`to12("1124")`は`"11:24 AM"`を返すべきです。

```js
assert.equal(to12("1124"), "11:24 AM");
```

`to12("0900")`は`"9:00 AM"`を返すべきです。

```js
assert.equal(to12("0900"), "9:00 AM");
```

`to12("1455")`は`"2:55 PM"`を返すべきです。

```js
assert.equal(to12("1455"), "2:55 PM");
```

`to12("2346")`は`"11:46 PM"`を返すべきです。

```js
assert.equal(to12("2346"), "11:46 PM");
```

`to12("0030")`は`"12:30 AM"`を返すべきです。

```js
assert.equal(to12("0030"), "12:30 AM");
```

# --seed--

## --seed-contents--

```js
function to12(time) {

  return time;
}
```

# --solutions--

```js
function convertHours(hours) {
  if (hours === 0) return 12;
  if (hours > 12) return hours - 12;
  return hours;
}

function to12(time) {
  const hours = parseInt(time.slice(0, 2), 10);
  const minutes = time.slice(2);
  const period = hours < 12 ? 'AM' : 'PM';

  return `${convertHours(hours)}:${minutes} ${period}`;
}
```
