---
id: 68cae5b538ff798bbd4da007
title: "チャレンジ 69: 片方だけの靴下"
challengeType: 28
dashedName: challenge-69
---

# --description--

最初に持っていた靴下のペア数を表す整数と、洗濯を何回行ったかを表す別の整数が与えられたとき、以下の条件に従って現在持っている完全な靴下のペア数を返してください。

- 洗濯を2回行うごとに、靴下が1足失われます。
- 洗濯を3回行うごとに、失くしていた靴下が1足見つかります。
- 洗濯を5回行うごとに、靴下が1足擦り切れて捨てなければなりません。
- 洗濯を10回行うごとに、靴下のペアを1組買います。
- 靴下の合計数が0未満になることはありません。
- ルールは重複する場合があります。例えば、洗濯10回目では、靴下を1足失い、1足捨て、新しいペアを1組買います。
- 完全な靴下のペア数を返してください。

# --hints--

`sockPairs(2, 5)` は `1` を返すべきです。

```js
assert.equal(sockPairs(2, 5), 1);
```

`sockPairs(1, 2)` は `0` を返すべきです。

```js
assert.equal(sockPairs(1, 2), 0);
```

`sockPairs(5, 11)` は `4` を返すべきです。

```js
assert.equal(sockPairs(5, 11), 4);
```

`sockPairs(6, 25)` は `3` を返すべきです。

```js
assert.equal(sockPairs(6, 25), 3);
```

`sockPairs(1, 8)` は `0` を返すべきです。

```js
assert.equal(sockPairs(1, 8), 0);
```

# --seed--

## --seed-contents--

```js
function sockPairs(pairs, cycles) {

  return pairs;
}
```

# --solutions--

```js
function sockPairs(pairs, cycles) {
  let socks = pairs * 2;

  for (let i = 1; i <= cycles; i++) {
    if (i % 2 === 0) socks -= 1;
    if (i % 3 === 0) socks += 1;
    if (i % 5 === 0) socks -= 1;
    if (i % 10 === 0) socks += 2;

    if (socks < 0) socks = 0;
  }

  return Math.floor(socks / 2);
}
```
