---
id: 68cae5b538ff798bbd4da005
title: "チャレンジ 67: メールアドレスバリデーター"
challengeType: 28
dashedName: challenge-67
---

# --description--

文字列が有効なメールアドレスかどうか、以下の条件に従って判定してください。

- `@` 記号がちょうど1つ含まれている必要があります。
- ローカル部分（`@` の前）：
  - 文字（`a-z`、`A-Z`）、数字（`0-9`）、ドット（`.`）、アンダースコア（`_`）、ハイフン（`-`）のみを含めることができます。
  - ドットで始まったり終わったりしてはいけません。
- ドメイン部分（`@` の後）：
  - 少なくとも1つのドットを含む必要があります。
  - ドットの後に少なくとも2文字の文字が続く形で終わる必要があります。
- ローカル部分もドメイン部分も、連続した2つのドットを含んではいけません。

# --hints--

`validate("a@b.cd")` は `true` を返すべきです。

```js
assert.isTrue(validate("a@b.cd"));
```

`validate("hell.-w.rld@example.com")` は `true` を返すべきです。

```js
assert.isTrue(validate("hell.-w.rld@example.com"));
```

`validate(".b@sh.rc")` は `false` を返すべきです。

```js
assert.isFalse(validate(".b@sh.rc"));
```

`validate("example@test.c0")` は `false` を返すべきです。

```js
assert.isFalse(validate("example@test.c0"));
```

`validate("freecodecamp.org")` は `false` を返すべきです。

```js
assert.isFalse(validate("freecodecamp.org"));
```

`validate("develop.ment_user@c0D!NG.R.CKS")` は `true` を返すべきです。

```js
assert.isTrue(validate("develop.ment_user@c0D!NG.R.CKS"));
```

`validate("hello.@wo.rld")` は `false` を返すべきです。

```js
assert.isFalse(validate("hello.@wo.rld"));
```

`validate("hello@world..com")` は `false` を返すべきです。

```js
assert.isFalse(validate("hello@world..com"));
```

`validate("develop..ment_user@c0D!NG.R.CKS")` は `false` を返すべきです。

```js
assert.isFalse(validate("develop..ment_user@c0D!NG.R.CKS"));
```

`validate("git@commit@push.io")` は `false` を返すべきです。

```js
assert.isFalse(validate("git@commit@push.io"));
```

# --seed--

## --seed-contents--

```js
function validate(email) {

  return email;
}
```

# --solutions--

```js
function validate(email) {
  if (email.includes('..')) return false;

  const parts = email.split('@');
  if (parts.length !== 2) return false;

  const [local, domain] = parts;

  if (local.startsWith('.') || local.endsWith('.')) return false;
  if (!/^[a-zA-Z0-9._-]+$/.test(local)) return false;
  if (!domain.includes('.')) return false;

  const tld = domain.split('.').pop();
  if (tld.length < 2 || !/^[a-zA-Z]+$/.test(tld)) return false;

  return true;
}
```
