---
id: 68cae5b538ff798bbd4da008
title: "チャレンジ 70: HTML 属性抽出器"
challengeType: 28
dashedName: challenge-70
---

# --description--

有効なHTML要素の文字列が与えられたとき、次の条件に従って要素の属性を返してください。

- 与えられるのは1つの要素だけです。
- 属性は `attribute="value"` の形式で与えられます。
- 各属性のプロパティと値をカンマで区切り、`["attribute1, value1", "attribute2, value2"]` の形式で文字列の配列として返してください。
- 属性は与えられた順番で返してください。
- 属性が見つからなければ、空の配列を返してください。

# --hints--

`extractAttributes('<span class="red"></span>')` は `["class, red"]` を返すべきです。

```js
assert.deepEqual(extractAttributes('<span class="red"></span>'), ["class, red"]);
```

`extractAttributes('<meta charset="UTF-8" />')` は `["charset, UTF-8"]` を返すべきです。

```js
assert.deepEqual(extractAttributes('<meta charset="UTF-8" />'), ["charset, UTF-8"]);
```

`extractAttributes("<p>Lorem ipsum dolor sit amet</p>")` は `[]` を返すべきです。

```js
assert.deepEqual(extractAttributes("<p>Lorem ipsum dolor sit amet</p>"), []);
```

`extractAttributes('<input name="email" type="email" required="true" />')` は `["name, email", "type, email", "required, true"]` を返すべきです。

```js
assert.deepEqual(extractAttributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"]);
```

`extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>')` は `["id, submit", "class, btn btn-primary"]` を返すべきです。

```js
assert.deepEqual(extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"]);
```

# --seed--

## --seed-contents--

```js
function extractAttributes(element) {

  return element;
}
```

# --solutions--

```js
function extractAttributes(element) {
  const regex = /(\w[\w-]*)="([^"]*)"/g;
  const result = [];
  let match;

  while ((match = regex.exec(element)) !== null) {
    const attr = match[1];
    const val = match[2];
    result.push(`${attr}, ${val}`);
  }

  return result;
}
```
