---
id: 68cae5b538ff798bbd4da00a
title: "チャレンジ 72: サーモスタット調整器 2"
challengeType: 28
dashedName: challenge-72
---

# --description--

部屋の現在の温度（華氏）と目標温度（摂氏）が与えられたとき、以下の条件に基づいて部屋の温度を調整する方法を示す文字列を返してください。

- 現在の温度が目標より低い場合は `"Heat: X degrees Fahrenheit"` を返してください。`X` は目標に達するために部屋を暖める必要がある華氏の度数で、小数点以下1桁に丸めます。
- 現在の温度が目標より高い場合は `"Cool: X degrees Fahrenheit"` を返してください。`X` は目標に達するために部屋を冷やす必要がある華氏の度数で、小数点以下1桁に丸めます。
- 現在の温度が目標と等しい場合は `"Hold"` を返してください。

摂氏を華氏に変換するには、摂氏温度に1.8を乗算し、結果に32を加えます（`F = (C * 1.8) + 32`）。

# --hints--

`adjustThermostat(32, 0)` は `"Hold"` を返すべきです。

```js
assert.equal(adjustThermostat(32, 0), "Hold");
```

`adjustThermostat(70, 25)` は `"Heat: 7.0 degrees Fahrenheit"` を返すべきです。

```js
assert.equal(adjustThermostat(70, 25), "Heat: 7.0 degrees Fahrenheit");
```

`adjustThermostat(72, 18)` は `"Cool: 7.6 degrees Fahrenheit"` を返すべきです。

```js
assert.equal(adjustThermostat(72, 18), "Cool: 7.6 degrees Fahrenheit");
```

`adjustThermostat(212, 100)` は `"Hold"` を返すべきです。

```js
assert.equal(adjustThermostat(212, 100), "Hold");
```

`adjustThermostat(59, 22)` は `"Heat: 12.6 degrees Fahrenheit"` を返すべきです。

```js
assert.equal(adjustThermostat(59, 22), "Heat: 12.6 degrees Fahrenheit");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(currentF, targetC) {

  return currentF;
}
```

# --solutions--

```js
function adjustThermostat(currentF, targetC) {
  const targetF = targetC * 1.8 + 32;
  const diff = Math.abs(targetF - currentF).toFixed(1);

  if (currentF < targetF) {
    return `Heat: ${diff} degrees Fahrenheit`;
  } else if (currentF > targetF) {
    return `Cool: ${diff} degrees Fahrenheit`;
  } else {
    return "Hold"
  }
}
```
