---
id: 68d2ba1468508398389487cf
title: "チャレンジ 74: お気に入りの曲"
challengeType: 28
dashedName: challenge-74
---

# --description--

iPodを覚えていますか？最初のモデルは24年前の今日、2001年10月23日に発売されました。

iPodのプレイリストを表す曲のオブジェクトの配列が与えられたとき、再生回数が最も多い2曲のタイトルを、再生回数が多い順に並べた配列を返してください。

- 各オブジェクトは`"title"`プロパティ（文字列）と`"plays"`プロパティ（整数）を持ちます。

# --hints--

`favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ])`は`["Sync or Swim", "Earbud Blues"]`を返すべきです。

```js
assert.deepEqual(favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ]), ["Sync or Swim", "Earbud Blues"]);
```

`favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ])`は`["Clickwheel Love", "99 Downloads"]`を返すべきです。

```js
assert.deepEqual(favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ]), ["Clickwheel Love", "99 Downloads"]);
```

`favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ])`は`["Song B", "Song C"]`を返すべきです。

```js
assert.deepEqual(favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ]), ["Song B", "Song C"]);
```

# --seed--

## --seed-contents--

```js
function favoriteSongs(playlist) {

  return playlist;
}
```

# --solutions--

```js
function favoriteSongs(playlist) {
  const sorted = [...playlist].sort((a, b) => b.plays - a.plays);
  return sorted.slice(0, 2).map(song => song.title);
}
```
