---
id: 68d2ba1468508398389487d0
title: "チャレンジ 75: 隠された宝物"
challengeType: 28
dashedName: challenge-75
---

# --description--

隠された宝物を含む海底の地図を表す2D配列と、宝探しの次の潜水地点の座標（`[row, column]`）を示す配列が与えられたとき、以下のルールに従って`"Empty"`、`"Found"`、または`"Recovered"`を返してください。

- 与えられた2D配列には、未回収の宝物がちょうど1つあり、それは複数のセルにまたがっています。
- 2D配列の各セルには次のいずれかの値が含まれます：
  - `"-"`：宝物なし。
  - `"O"`：まだ見つかっていない宝物の一部。
  - `"X"`：すでに見つかった宝物の一部。

- 潜水地点に宝物がなければ、`"Empty"`を返してください。
- 潜水地点で宝物が見つかったが、他にまだ見つかっていない宝物の部分がある場合は、`"Found"`を返してください。
- 潜水地点で宝物の最後の未発見部分が見つかった場合は、`"Recovered"`を返してください。

例えば、次のような場合：

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

潜水地点の座標が`[2, 1]`であれば、潜水で宝物の最後の未発見部分が見つかったため、`"Recovered"`を返します。

# --hints--

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])`は`"Recovered"`を返すべきです。

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Recovered");
```

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])`は`"Empty"`を返すべきです。

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Empty");
```

`dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])`は`"Found"`を返すべきです。

```js
assert.equal(dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Found");
```

`dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])`は`"Found"`を返すべきです。

```js
assert.equal(dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Found");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])`は`"Recovered"`を返すべきです。

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Recovered");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])`は`"Empty"`を返すべきです。

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Empty");
```

# --seed--

## --seed-contents--

```js
function dive(map, coordinates) {

  return map;
}
```

# --solutions--

```js
function dive(map, coordinates) {
  const [row, col] = coordinates;
  const target = map[row][col];

  if (target === "-") {
    return "Empty";
  }

  if (target === "O") {
    for (let r = 0; r < map.length; r++) {
      for (let c = 0; c < map[r].length; c++) {
        if ((r !== row || c !== col) && map[r][c] === "O") {
          return "Found";
        }
      }
    }
    return "Recovered";
  }

  return "Found";
}
```
