---
id: 68d30845cc08266018fc46bc
title: "チャレンジ 76: 補色の DNA"
challengeType: 28
dashedName: challenge-76
---

# --description--

DNA シーケンスを表す文字列が与えられたとき、次のルールに従ってその補色の鎖を返してください。

- DNA は `"A"`、`"C"`、`"G"`、`"T"` の文字で構成されています。
- `"A"` と `"T"` の文字は互いに補色の関係です。
- `"C"` と `"G"` の文字は互いに補色の関係です。

たとえば、`"ACGT"` が与えられたら、`"TGCA"` を返してください。

# --hints--

`complementaryDNA("ACGT")` は `"TGCA"` を返すべきです。

```js
assert.equal(complementaryDNA("ACGT"), "TGCA");
```

`complementaryDNA("ATGCGTACGTTAGC")` は `"TACGCATGCAATCG"` を返すべきです。

```js
assert.equal(complementaryDNA("ATGCGTACGTTAGC"), "TACGCATGCAATCG");
```

`complementaryDNA("GGCTTACGATCGAAG")` は `"CCGAATGCTAGCTTC"` を返すべきです。

```js
assert.equal(complementaryDNA("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC");
```

`complementaryDNA("GATCTAGCTAGGCTAGCTAG")` は `"CTAGATCGATCCGATCGATC"` を返すべきです。

```js
assert.equal(complementaryDNA("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC");
```

# --seed--

## --seed-contents--

```js
function complementaryDNA(strand) {

  return strand;
}
```

# --solutions--

```js
function complementaryDNA(strand) {
  const complements = { A: "T", T: "A", C: "G", G: "C" };
  return strand
    .split("")
    .map(c => complements[c])
    .join("");
}
```
