---
id: 68d30845cc08266018fc46bd
title: "チャレンジ 77: 期間フォーマッター"
challengeType: 28
dashedName: challenge-77
---

# --description--

整数の秒数が与えられたとき、同じ期間を `"H:MM:SS"` 形式で表す文字列を返してください。ここで `"H"` は時間数、`"MM"` は分数、`"SS"` は秒数です。次のルールに従って時間を返してください。

- 秒数：常に2桁で表示してください。
- 分数：必要ない場合は先頭のゼロを省略してください。期間が1分未満の場合は `"0"` を使用してください。
- 時間：0より大きい場合のみ含めてください。

# --hints--

`format(500)` は `"8:20"` を返すべきです。

```js
assert.equal(format(500), "8:20");
```

`format(4000)` は `"1:06:40"` を返すべきです。

```js
assert.equal(format(4000), "1:06:40");
```

`format(1)` は `"0:01"` を返すべきです。

```js
assert.equal(format(1), "0:01");
```

`format(5555)` は `"1:32:35"` を返すべきです。

```js
assert.equal(format(5555), "1:32:35");
```

`format(99999)` は `"27:46:39"` を返すべきです。

```js
assert.equal(format(99999), "27:46:39");
```

# --seed--

## --seed-contents--

```js
function format(seconds) {

  return seconds;
}
```

# --solutions--

```js
function format(seconds) {
  const h = Math.floor(seconds / 3600);
  const m = Math.floor((seconds % 3600) / 60);
  const s = seconds % 60;

  const secondsStr = s.toString().padStart(2, "0");
  const minutesStr = m.toString();

  if (h > 0) {
    const hoursStr = h.toString();
    return `${hoursStr}:${minutesStr.padStart(2, "0")}:${secondsStr}`;
  } else {
    return `${minutesStr}:${secondsStr}`;
  }
}
```
