---
id: 68d30845cc08266018fc46be
title: "チャレンジ 78: 整数シーケンス"
challengeType: 28
dashedName: challenge-78
---

# --description--

正の整数が与えられたら、`1`からその数までのすべての整数を数値の順に並べた文字列を返してください。

例えば、`5`が与えられたら、`"12345"`を返します。

# --hints--

`sequence(5)`は`"12345"`を返すはずです。

```js
assert.equal(sequence(5), "12345");
```

`sequence(10)`は`"12345678910"`を返すはずです。

```js
assert.equal(sequence(10), "12345678910");
```

`sequence(1)`は`"1"`を返すはずです。

```js
assert.strictEqual(sequence(1), "1");
```

`sequence(27)`は`"123456789101112131415161718192021222324252627"`を返すはずです。

```js
assert.equal(sequence(27), "123456789101112131415161718192021222324252627");
```

# --seed--

## --seed-contents--

```js
function sequence(n) {

  return n;
}
```

# --solutions--

```js
function sequence(n) {
  let result = "";
  for (let i = 1; i <= n; i++) {
    result += i;
  }
  return result;
}
```
