---
id: 68e39ed6106dac2f0a98fd65
title: "チャレンジ 83: 署名のバリデーション"
challengeType: 28
dashedName: challenge-83
---

# --description--

メッセージの文字列、秘密鍵の文字列、署名の数値が与えられたとき、このエンコード方法を使って署名が有効かどうかを判定してください。

- メッセージと秘密鍵の文字は次の値を持ちます：
  - `a` から `z` はそれぞれ `1` から `26` の値を持ちます。
  - `A` から `Z` はそれぞれ `27` から `52` の値を持ちます。
- その他の文字は値を持ちません。
- 署名はメッセージの合計値と秘密鍵の合計値を足したものとして計算します。

例えば、メッセージが `"foo"`、秘密鍵が `"bar"` の場合、署名は `57` になります：

```md
f (6) + o (15) + o (15) = 36
b (2) + a (1) + r (18) = 21
36 + 21 = 57
```

計算した署名が与えられた署名と一致するかどうかを確認してください。

# --hints--

`verify("foo", "bar", 57)` は `true` を返すはずです。

```js
assert.isTrue(verify("foo", "bar", 57));
```

`verify("foo", "bar", 54)` は `false` を返すはずです。

```js
assert.isFalse(verify("foo", "bar", 54));
```

`verify("freeCodeCamp", "Rocks", 238)` は `true` を返すはずです。

```js
assert.isTrue(verify("freeCodeCamp", "Rocks", 238));
```

`verify("Is this valid?", "No", 210)` は `false` を返すはずです。

```js
assert.isFalse(verify("Is this valid?", "No", 210));
```

`verify("Is this valid?", "Yes", 233)` は `true` を返すはずです。

```js
assert.isTrue(verify("Is this valid?", "Yes", 233));
```

`verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514)` は `true` を返すはずです。

```js
assert.isTrue(verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514));
```

# --seed--

## --seed-contents--

```js
function verify(message, key, signature) {

  return message;
}
```

# --solutions--

```js
function verify(message, key, signature) {
  function charValue(ch) {
    if (ch >= 'a' && ch <= 'z') return ch.charCodeAt(0) - 'a'.charCodeAt(0) + 1;
    if (ch >= 'A' && ch <= 'Z') return ch.charCodeAt(0) - 'A'.charCodeAt(0) + 27;
    return 0;
  }

  function computeSum(str) {
    let sum = 0;
    for (let ch of str) {
      sum += charValue(ch);
    }
    return sum;
  }

  const total = computeSum(message) + computeSum(key);
  return total === signature;
}
```
