---
id: 68e39ed6106dac2f0a98fd64
title: "チャレンジ 82: SpOoKy~CaSe"
challengeType: 28
dashedName: challenge-82
---

# --description--

変数名を表す文字列が与えられたとき、次の条件に従って「spooky case」に変換してください。

- すべてのアンダースコア（`_`）とハイフン（`-`）をチルダ（`~`）に置き換えます。
- 文字列の最初の文字と、その後の文字のうちチルダを除いた偶数番目の文字を大文字にします。その他の文字はすべて小文字にします。

例えば、`hello_world`が与えられたら、`HeLlO~wOrLd`を返してください。

# --hints--

`spookify("hello_world")`は`"HeLlO~wOrLd"`を返すべきです。

```js
assert.equal(spookify("hello_world"), "HeLlO~wOrLd");
```

`spookify("Spooky_Case")`は`"SpOoKy~CaSe"`を返すべきです。

```js
assert.equal(spookify("Spooky_Case"), "SpOoKy~CaSe");
```

`spookify("TRICK-or-TREAT")`は`"TrIcK~oR~tReAt"`を返すべきです。

```js
assert.equal(spookify("TRICK-or-TREAT"), "TrIcK~oR~tReAt");
```

`spookify("c_a-n_d-y_-b-o_w_l")`は`"C~a~N~d~Y~~b~O~w~L"`を返すべきです。

```js
assert.equal(spookify("c_a-n_d-y_-b-o_w_l"), "C~a~N~d~Y~~b~O~w~L");
```

`spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts")`は`"ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS"`を返すべきです。

```js
assert.equal(spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts"), "ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS");
```

# --seed--

## --seed-contents--

```js
function spookify(boo) {

  return boo;
}
```

# --solutions--

```js
function spookify(boo) {
  const replaced = boo.replace(/[_-]/g, "~");
  
  let result = "";
  let capitalize = true;

  for (let char of replaced) {
    if (char === "~") {
      result += char;
    } else {
      result += capitalize ? char.toUpperCase() : char.toLowerCase();
      capitalize = !capitalize;
    }
  }

  return result;
}
```
