---
id: 68e39ed6106dac2f0a98fd66
title: "チャレンジ 84: 感染"
challengeType: 28
dashedName: challenge-84
---

# --description--

1988年11月2日、最初の大規模なインターネットワームがリリースされ、わずか1日でインターネットに接続されたコンピューターの約10％が感染しました。

このチャレンジでは、インターネットワームがリリースされてから経過した日数が与えられ、そのルールに従って感染したコンピューターの数を求めてください。

- 0日目に最初のコンピューターが感染します。
- その後の日は、感染したコンピューターの数が倍増します。
- 3日ごとに、ウイルスが広がった後にパッチが適用され、感染したコンピューターの数が20％減少します。パッチで減少するコンピューターの数は小数点以下を切り上げてください。

例えば、

- 0日目：合計1台のコンピューターが感染します。
- 1日目：合計2台のコンピューターが感染します。
- 2日目：合計4台のコンピューターが感染します。
- 3日目：合計8台のコンピューターが感染します。その後パッチを適用します。8台 × 20％ = 1.6台がパッチ適用。1.6を切り上げて2台。8台 - 2台 = 6台が3日目終了時の感染数です。

与えられた日数が経過した後の感染したコンピューターの合計数を返してください。

# --hints--

`infected(1)` は `2` を返すべきです。

```js
assert.equal(infected(1), 2);
```

`infected(3)` は `6` を返すべきです。

```js
assert.equal(infected(3), 6);
```

`infected(8)` は `152` を返すべきです。

```js
assert.equal(infected(8), 152);
```

`infected(17)` は `39808` を返すべきです。

```js
assert.equal(infected(17), 39808);
```

`infected(25)` は `5217638` を返すべきです。

```js
assert.equal(infected(25), 5217638);
```

# --seed--

## --seed-contents--

```js
function infected(days) {

  return days;
}
```

# --solutions--

```js
function infected(days) {
  let infected = 1;

  for (let day = 1; day <= days; day++) {
    infected *= 2;

    if (day % 3 === 0) {
      let patched = Math.ceil(infected * 0.2);
      infected -= patched;
    }
  }

  return infected;
}
```
