---
id: 68e39ed6106dac2f0a98fd62
title: "チャレンジ 80: メールソーター"
challengeType: 28
dashedName: challenge-80
---

# --description--

1971年10月29日に、最初のメールが送信され、現在も使われている`username@domain`形式が導入されました。現在では、数十億のメールアドレスがあります。

このチャレンジでは、メールアドレスのリストが与えられ、ドメイン名（`@`の後の部分）でアルファベット順に、次にユーザー名（`@`の前の部分）でソートする必要があります。

- ソートは大文字・小文字を区別しません。
- 同じドメインのメールが複数ある場合は、ユーザー名でソートしてください。
- ソートされたアドレスの配列を返してください。
- 返すアドレスは元の大文字・小文字を保持してください。

例えば、`["jill@mail.com", "john@example.com", "jane@example.com"]`が与えられた場合、`["jane@example.com", "john@example.com", "jill@mail.com"]`を返してください。

# --hints--

`sort(["jill@mail.com", "john@example.com", "jane@example.com"])`は`["jane@example.com", "john@example.com", "jill@mail.com"]`を返すべきです。

```js
assert.deepEqual(sort(["jill@mail.com", "john@example.com", "jane@example.com"]), ["jane@example.com", "john@example.com", "jill@mail.com"]);
```

`sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"])`は`["bob@mail.com", "carol@mail.com", "alice@zoo.com"]`を返すべきです。

```js
assert.deepEqual(sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"]), ["bob@mail.com", "carol@mail.com", "alice@zoo.com"]);
```

`sort(["user@z.com", "user@y.com", "user@x.com"])`は`["user@x.com", "user@y.com", "user@z.com"]`を返すべきです。

```js
assert.deepEqual(sort(["user@z.com", "user@y.com", "user@x.com"]), ["user@x.com", "user@y.com", "user@z.com"]);
```

`sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"])`は`["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]`を返すべきです。

```js
assert.deepEqual(sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"]), ["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]);
```

`sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"])`は`["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]`を返すべきです。

```js
assert.deepEqual(sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"]), ["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]);
```

# --seed--

## --seed-contents--

```js
function sort(emails) {

  return emails;
}
```

# --solutions--

```js
function sort(emails) {
  return emails.slice().sort((a, b) => {
    const [userA, domainA] = a.split('@');
    const [userB, domainB] = b.split('@');

    const domainCompare = domainA.toLowerCase().localeCompare(domainB.toLowerCase());
    if (domainCompare !== 0) return domainCompare;

    return userA.toLowerCase().localeCompare(userB.toLowerCase());
  });
}
```
