---
id: 68e39ed6106dac2f0a98fd63
title: "チャレンジ 81: 第N番目の素数"
challengeType: 28
dashedName: challenge-81
---

# --description--

素数とは、1より大きい正の整数で、1とその数自身以外に割り切れる数がないものです。最初の5つの素数は `2`、`3`、`5`、`7`、および `11` です。

正の整数 `n` が与えられたとき、`n` 番目の素数を返してください。例えば、`5` が与えられた場合、5番目の素数である `11` を返します。

# --hints--

`nthPrime(5)` は `11` を返すべきです。

```js
assert.equal(nthPrime(5), 11);
```

`nthPrime(10)` は `29` を返すべきです。

```js
assert.equal(nthPrime(10), 29);
```

`nthPrime(16)` は `53` を返すべきです。

```js
assert.equal(nthPrime(16), 53);
```

`nthPrime(99)` は `523` を返すべきです。

```js
assert.equal(nthPrime(99), 523);
```

`nthPrime(1000)` は `7919` を返すべきです。

```js
assert.equal(nthPrime(1000), 7919);
```

# --seed--

## --seed-contents--

```js
function nthPrime(n) {

  return n;
}
```

# --solutions--

```js
function nthPrime(n) {
  const primes = [];
  let num = 2;

  while (primes.length < n) {
    let isPrime = true;
    for (let i = 2; i * i <= num; i++) {
      if (num % i === 0) {
        isPrime = false;
        break;
      }
    }
    if (isPrime) primes.push(num);
    num++;
  }

  return primes[n - 1];
}
```
