---
id: 68ee9e3066cfd4eb2328e8a8
title: "チャレンジ 89: カードの組み合わせを数える"
challengeType: 28
dashedName: challenge-89
---

# --description--

標準的なトランプのデッキには、各スートに13枚のユニークなカードがあります。デッキから選ぶカードの枚数を表す整数が与えられたとき、選べるユニークなカードの組み合わせの数を返してください。

- 順序は関係ありません。カードAを選んでからカードBを選ぶのは、カードBを選んでからカードAを選ぶのと同じです。

例えば、`52`が与えられたら`1`を返してください。52枚のデッキから52枚選ぶ組み合わせは1通りだけです。また、`2`が与えられたら`1326`を返してください。デッキから2枚選ぶと1326通りのカードの組み合わせが考えられます。

# --hints--

`combinations(52)`は`1`を返すべきです。

```js
assert.equal(combinations(52), 1);
```

`combinations(1)`は`52`を返すべきです。

```js
assert.equal(combinations(1), 52);
```

`combinations(2)`は`1326`を返すべきです。

```js
assert.equal(combinations(2), 1326);
```

`combinations(5)`は`2598960`を返すべきです。

```js
assert.equal(combinations(5), 2598960);
```

`combinations(10)`は`15820024220`を返すべきです。

```js
assert.equal(combinations(10), 15820024220);
```

`combinations(50)`は`1326`を返すべきです。

```js
assert.equal(combinations(50), 1326);
```

# --seed--

## --seed-contents--

```js
function combinations(cards) {

  return cards;
}
```

# --solutions--

```js
function combinations(cards) {
  const n = 52;

  function factorial(x) {
    let result = 1;
    for (let i = 2; i <= x; i++) {
      result *= i;
    }
    return result;
  }

  return factorial(n) / (factorial(cards) * factorial(n - cards));
}
```
