---
id: 68ee9e3066cfd4eb2328e8a4
title: "チャレンジ 85: 単語カウンター"
challengeType: 28
dashedName: challenge-85
---

# --description--

文章の文字列が与えられたとき、その文章に含まれる単語の数を返してください。

- 単語とは、空白以外の文字のシーケンスであり、単語同士は1つの空白で区切られています。

# --hints--

`countWords("Hello world")` は `2` を返すべきです。

```js
assert.equal(countWords("Hello world"), 2);
```

`countWords("The quick brown fox jumps over the lazy dog.")` は `9` を返すべきです。

```js
assert.equal(countWords("The quick brown fox jumps over the lazy dog."), 9);
```

`countWords("I like coding challenges!")` は `4` を返すべきです。

```js
assert.equal(countWords("I like coding challenges!"), 4);
```

`countWords("Complete the challenge in JavaScript and Python.")` は `7` を返すべきです。

```js
assert.equal(countWords("Complete the challenge in JavaScript and Python."), 7);
```

`countWords("The missing semi-colon crashed the entire internet.")` は `7` を返すべきです。

```js
assert.equal(countWords("The missing semi-colon crashed the entire internet."), 7);
```

# --seed--

## --seed-contents--

```js
function countWords(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countWords(sentence) {

  return sentence.split(' ').length;
}
```
