---
id: 68ee9e3066cfd4eb2328e8a6
title: "チャレンジ 87: マトリックスビルダー"
challengeType: 28
dashedName: challenge-87
---

# --description--

2つの整数（行数と列数）が与えられたとき、そのサイズのゼロ（`0`）で埋められたマトリックス（配列の配列）を返してください。

たとえば、`2` と `3` が与えられた場合、次を返します：

```json
[
  [0, 0, 0],
  [0, 0, 0]
]
```

# --hints--

`buildMatrix(2, 3)` は `[[0, 0, 0], [0, 0, 0]]` を返すはずです。

```js
assert.deepEqual(buildMatrix(2, 3), [[0, 0, 0], [0, 0, 0]]);
```

`buildMatrix(3, 2)` は `[[0, 0], [0, 0], [0, 0]]` を返すはずです。

```js
assert.deepEqual(buildMatrix(3, 2), [[0, 0], [0, 0], [0, 0]]);
```

`buildMatrix(4, 3)` は `[[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]` を返すはずです。

```js
assert.deepEqual(buildMatrix(4, 3), [[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]);
```

`buildMatrix(9, 1)` は `[[0], [0], [0], [0], [0], [0], [0], [0], [0]]` を返すはずです。

```js
assert.deepEqual(buildMatrix(9, 1), [[0], [0], [0], [0], [0], [0], [0], [0], [0]]);
```

# --seed--

## --seed-contents--

```js
function buildMatrix(rows, cols) {

  return rows;
}
```

# --solutions--

```js
function buildMatrix(rows, cols) {
  const matrix = [];
  for (let i = 0; i < rows; i++) {
    const row = new Array(cols).fill(0);
    matrix.push(row);
  }
  return matrix;
}
```
