---
id: 68ee9e3066cfd4eb2328e8a7
title: "チャレンジ 88: 曜日を見つける"
challengeType: 28
dashedName: challenge-88
---

# --description--

`YYYY-MM-DD` 形式の文字列の日付が与えられたら、曜日を返してください。

有効な返却曜日は以下の通りです：

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

タイムゾーンは無視してください。

# --hints--

`getWeekday("2025-11-06")` は `Thursday` を返すはずです。

```js
assert.equal(getWeekday("2025-11-06"), "Thursday");
```

`getWeekday("1999-12-31")` は `Friday` を返すはずです。

```js
assert.equal(getWeekday("1999-12-31"), "Friday");
```

`getWeekday("1111-11-11")` は `Saturday` を返すはずです。

```js
assert.equal(getWeekday("1111-11-11"), "Saturday");
```

`getWeekday("2112-12-21")` は `Wednesday` を返すはずです。

```js
assert.equal(getWeekday("2112-12-21"), "Wednesday");
```

`getWeekday("2345-10-01")` は `Monday` を返すはずです。

```js
assert.equal(getWeekday("2345-10-01"), "Monday");
```

# --seed--

## --seed-contents--

```js
function getWeekday(dateString) {

  return dateString;
}
```

# --solutions--

```js
function getWeekday(dateString) {
  const days = ["Sunday", "Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday"];
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  return days[date.getUTCDay()];
}
```
