---
id: 68ee9e3066cfd4eb2328e8a5
title: "チャレンジ 86: 画像検索"
challengeType: 28
dashedName: challenge-86
---

# --description--

2001年11月4日、Googleは画像検索を開始し、検索語句を使って画像を見つけられるようになりました。このチャレンジでは、その画像検索を模倣します。

画像名の配列と検索語句が与えられたとき、検索語句を含む画像名の配列を返してください。

- 検索語句の一致は大文字・小文字を区別しません。
- 画像は入力配列に現れる順番のまま返してください。

# --hints--

`imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog")` は `["dog.png"]` を返すべきです。

```js
assert.deepEqual(imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog"), ["dog.png"]);
```

`imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun")` は `["Sunset.jpg", "sunflower.jpeg"]` を返すべきです。

```js
assert.deepEqual(imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun"), ["Sunset.jpg", "sunflower.jpeg"]);
```

`imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG")` は `["Moon.png", "stars.png"]` を返すべきです。

```js
assert.deepEqual(imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG"), ["Moon.png", "stars.png"]);
```

`imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat")` は `["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]` を返すべきです。

```js
assert.deepEqual(imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat"), ["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]);
```

# --seed--

## --seed-contents--

```js
function imageSearch(images, term) {

  return images;
}
```

# --solutions--

```js
function imageSearch(images, term) {
  const lowerTerm = term.toLowerCase();
  return images.filter(img => img.toLowerCase().includes(lowerTerm));
}
```
