---
id: 68f6587287ad1f4ad39b0c85
title: "チャレンジ 99: 指紋テスト"
challengeType: 28
dashedName: challenge-99
---

# --description--

2つの指紋を表す文字列が与えられたとき、次のルールに従って一致しているかどうかを判定してください。

- 各指紋は小文字の文字のみで構成されます（`a-z`）。
- 2つの指紋が一致するとみなされるのは以下の場合です：
  - 長さが同じであること。
  - 異なる文字の数が指紋の長さの10％を超えないこと。

# --hints--

`isMatch("helloworld", "helloworld")` は `true` を返すべきです。

```js
assert.isTrue(isMatch("helloworld", "helloworld"));
```

`isMatch("helloworld", "helloworlds")` は `false` を返すべきです。

```js
assert.isFalse(isMatch("helloworld", "helloworlds"));
```

`isMatch("helloworld", "jelloworld")` は `true` を返すべきです。

```js
assert.isTrue(isMatch("helloworld", "jelloworld"));
```

`isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog")` は `true` を返すべきです。

```js
assert.isTrue(isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog"));
```

`isMatch("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog")` は `true` を返すべきです。

```js
assert.isTrue(isMatch("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog"));
```

`isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat")` は `false` を返すべきです。

```js
assert.isFalse(isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat"));
```

# --seed--

## --seed-contents--

```js
function isMatch(fingerprintA, fingerprintB) {

  return fingerprintA;
}
```

# --solutions--

```js
function isMatch(fingerprintA, fingerprintB) {
  if (fingerprintA.length !== fingerprintB.length) return false;

  const length = fingerprintA.length;
  let mismatches = 0;

  for (let i = 0; i < length; i++) {
    if (fingerprintA[i] !== fingerprintB[i]) {
      mismatches++;
      if (mismatches > length * 0.1) return false;
    }
  }

  return true;
}
```
