---
id: 68f6587287ad1f4ad39b0c81
title: "チャレンジ 95: 配列のシフト"
challengeType: 28
dashedName: challenge-95
---

# --description--

配列と、配列を何ポジションシフトするかを表す整数が与えられたとき、シフト後の配列を返してください。

- 正の整数は配列を左にシフトします。
- 負の整数は配列を右にシフトします。
- シフトは配列の端をまたいで循環します。

たとえば、`[1, 2, 3]` と `1` が与えられた場合、配列を左に1シフトして `[2, 3, 1]` を返します。

# --hints--

`shiftArray([1, 2, 3], 1)` は `[2, 3, 1]` を返すはずです。

```js
assert.deepEqual(shiftArray([1, 2, 3], 1), [2, 3, 1]);
```

`shiftArray([1, 2, 3], -1)` は `[3, 1, 2]` を返すはずです。

```js
assert.deepEqual(shiftArray([1, 2, 3], -1), [3, 1, 2]);
```

`shiftArray(["alpha", "bravo", "charlie"], 5)` は `["charlie", "alpha", "bravo"]` を返すはずです。

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], 5), ["charlie", "alpha", "bravo"]);
```

`shiftArray(["alpha", "bravo", "charlie"], -11)` は `["bravo", "charlie", "alpha"]` を返すはずです。

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], -11), ["bravo", "charlie", "alpha"]);
```

`shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15)` は `[5, 6, 7, 8, 9, 0, 1, 2, 3, 4]` を返すはずです。

```js
assert.deepEqual(shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15), [5, 6, 7, 8, 9, 0, 1, 2, 3, 4]);
```

# --seed--

## --seed-contents--

```js
function shiftArray(arr, n) {

  return arr;
}
```

# --solutions--

```js
function shiftArray(arr, n) {
  const len = arr.length;
  n = n % len;
  if (n < 0) n += len;

  return arr.slice(n).concat(arr.slice(0, n));
}
```
