---
id: 68f6587287ad1f4ad39b0c80
title: "チャレンジ 94: メール署名ジェネレーター"
challengeType: 28
dashedName: challenge-94
---

# --description--

名前、役職、会社名の文字列が与えられたとき、以下のルールに従ってメール署名を1つの文字列として返してください。

- 名前は最初に表示し、名前の最初の文字に応じた接頭辞を付けます（大文字・小文字は区別しません）。
  - `A-I`で始まる名前には`>>`を接頭辞として使います。
  - `J-R`で始まる名前には`--`を接頭辞として使います。
  - `S-Z`で始まる名前には`::`を接頭辞として使います。
- 名前の後にはカンマとスペース（`, `）を付けます。
- 役職と会社名はカンマとスペースの後に続き、`" at "`（前後にスペースあり）で区切ります。

例えば、`"Quinn Waverly"`、`"Founder and CEO"`、`"TechCo"`が与えられた場合、`"--Quinn Waverly, Founder and CEO at TechCo"`を返します。

# --hints--

`generateSignature("Quinn Waverly", "Founder and CEO", "TechCo")`は`"--Quinn Waverly, Founder and CEO at TechCo"`を返すはずです。

```js
assert.equal(generateSignature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo");
```

`generateSignature("Alice Reed", "Engineer", "TechCo")`は`">>Alice Reed, Engineer at TechCo"`を返すはずです。

```js
assert.equal(generateSignature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo");
```

`generateSignature("Tina Vaughn", "Developer", "example.com")`は`"::Tina Vaughn, Developer at example.com"`を返すはずです。

```js
assert.equal(generateSignature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com");
```

`generateSignature("B. B.", "Product Tester", "AcmeCorp")`は`">>B. B., Product Tester at AcmeCorp"`を返すはずです。

```js
assert.equal(generateSignature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp");
```

`generateSignature("windstorm", "Cloud Architect", "Atmospheronics")`は`"::windstorm, Cloud Architect at Atmospheronics"`を返すはずです。

```js
assert.equal(generateSignature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics");
```

# --seed--

## --seed-contents--

```js
function generateSignature(name, title, company) {

  return name;
}
```

# --solutions--

```js
function generateSignature(name, title, company) {
  const firstLetter = name[0].toUpperCase();
  let prefix;
  if ("ABCDEFGHI".includes(firstLetter)) {
    prefix = ">>";
  } else if ("JKLMNOPQR".includes(firstLetter)) {
    prefix = "--";
  } else {
    prefix = "::";
  }

  return `${prefix}${name}, ${title} at ${company}`;
}
```
