---
id: 68f6587287ad1f4ad39b0c84
title: "チャレンジ 98: 長方形の数を数える"
challengeType: 28
dashedName: challenge-98
---

# --description--

長方形の幅と高さを表す2つの正の整数が与えられたとき、その長方形の中にいくつの長方形が収まるかを求めてください。

- 幅と高さが整数の長方形のみを数えてください。

例えば、`1` と `3` が与えられた場合、`6` を返します。1x1の長方形が3つ、1x2の長方形が2つ、1x3の長方形が1つです。

# --hints--

`countRectangles(1, 3)` は `6` を返すべきです。

```js
assert.equal(countRectangles(1, 3), 6);
```

`countRectangles(3, 2)` は `18` を返すべきです。

```js
assert.equal(countRectangles(3, 2), 18);
```

`countRectangles(1, 2)` は `3` を返すべきです。

```js
assert.equal(countRectangles(1, 2), 3);
```

`countRectangles(5, 4)` は `150` を返すべきです。

```js
assert.equal(countRectangles(5, 4), 150);
```

`countRectangles(11, 19)` は `12540` を返すべきです。

```js
assert.equal(countRectangles(11, 19), 12540);
```

# --seed--

## --seed-contents--

```js
function countRectangles(width, height) {

  return width;
}
```

# --solutions--

```js
function countRectangles(width, height) {
  let count = 0;

  for (let w = 1; w <= width; w++) {
    for (let h = 1; h <= height; h++) {
      count += (width - w + 1) * (height - h + 1);
    }
  }

  return count;
}
```
