---
id: 68f6587287ad1f4ad39b0c7c
title: "チャレンジ90：文字数制限"
challengeType: 28
dashedName: challenge-90
---

# --description--

このチャレンジでは、文字列が与えられ、ソーシャルメディアの投稿に収まるかどうかを判定します。以下のルールに基づいて、次の文字列を返してください：

- `"short post"` は40文字以内に収まる場合に返します。
- `"long post"` は40文字を超え、80文字以内に収まる場合に返します。
- `"invalid post"` はどちらの制限にも収まらない場合に返します。

# --hints--

`canPost("Hello world")` は `"short post"` を返すべきです。

```js
assert.equal(canPost("Hello world"), "short post");
```

`canPost("This is a longer message but still under eighty characters.")` は `"long post"` を返すべきです。

```js
assert.equal(canPost("This is a longer message but still under eighty characters."), "long post");
```

`canPost("This message is too long to fit into either of the character limits for a social media post.")` は `"invalid post"` を返すべきです。

```js
assert.equal(canPost("This message is too long to fit into either of the character limits for a social media post."), "invalid post");
```

# --seed--

## --seed-contents--

```js
function canPost(message) {

  return message;
}
```

# --solutions--

```js
function canPost(message) {
  if (message.length <= 40) {
    return "short post"
  } else if (message.length <= 80) {
    return "long post"
  } else {
    return "invalid post"
  }
}
```
