---
id: 68f6587287ad1f4ad39b0c83
title: "チャレンジ 97: 最大公約数"
challengeType: 28
dashedName: challenge-97
---

# --description--

2つの正の整数が与えられたとき、その最大公約数（GCD）を返してください。

- 2つの整数の最大公約数とは、両方の数を割り切ることができる最大の数のことです。余りは出ません。

例えば、`4` の約数は `1`、`2`、`4` です。`6` の約数は `1`、`2`、`3`、`6` です。したがって、`4` と `6` が与えられたら、両方の約数の集合に現れる最大の数である `2` を返してください。

# --hints--

`gcd(4, 6)` は `2` を返すべきです。

```js
assert.equal(gcd(4, 6), 2);
```

`gcd(20, 15)` は `5` を返すべきです。

```js
assert.equal(gcd(20, 15), 5);
```

`gcd(13, 17)` は `1` を返すべきです。

```js
assert.equal(gcd(13, 17), 1);
```

`gcd(654, 456)` は `6` を返すべきです。

```js
assert.equal(gcd(654, 456), 6);
```

`gcd(3456, 4320)` は `864` を返すべきです。

```js
assert.equal(gcd(3456, 4320), 864);
```

# --seed--

## --seed-contents--

```js
function gcd(x, y) {

  return x;
}
```

# --solutions--

```js
function gcd(x, y) {
  while (y !== 0) {
    const temp = y;
    y = x % y;
    x = temp;
  }
  return x;
}
```
