---
id: 68f6587287ad1f4ad39b0c7d
title: "チャレンジ 91: ワードサーチ"
challengeType: 28
dashedName: challenge-91
---

# --description--

1文字ずつの文字が入った行列（配列の配列）と探す単語が与えられたとき、その単語が行列内で始まる位置と終わる位置のインデックスを返してください。

- 与えられる行列はすべて小文字の文字で埋められています（`a-z`）。
- 探す単語は必ず行列内に1回だけ存在します。
- 探す単語は必ず次のいずれかの方向の直線上にあります：
  - 左から右
  - 右から左
  - 上から下
  - 下から上

例えば、次の行列が与えられた場合：

```md
[
  ["a", "c", "t"],
  ["t", "a", "t"],
  ["c", "t", "c"]
]
```

そして単語 `"cat"` を探すと、次を返します：

```md
[[0, 1], [2, 1]]
```

ここで `[0, 1]` は `"c"` のインデックス（単語の開始位置）、`[2, 1]` は `"t"` のインデックス（単語の終了位置）です。

# --hints--

`findWord([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat")` は `[[0, 1], [2, 1]]` を返すべきです。

```js
assert.deepEqual(findWord([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat"), [[0, 1], [2, 1]]);
```

`findWord([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog")` は `[[0, 0], [0, 2]]` を返すべきです。

```js
assert.deepEqual(findWord([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog"), [[0, 0], [0, 2]]);
```

`findWord([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish")` は `[[3, 3], [0, 3]]` を返すべきです。

```js
assert.deepEqual(findWord([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish"), [[3, 3], [0, 3]]);
```

`findWord([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox")` は `[[1, 3], [1, 1]]` を返すべきです。

```js
assert.deepEqual(findWord([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox"), [[1, 3], [1, 1]]);
```

# --seed--

## --seed-contents--

```js
function findWord(matrix, word) {

  return matrix;
}
```

# --solutions--

```js
function findWord(matrix, word) {
  const rows = matrix.length;
  const cols = matrix[0].length;
  const len = word.length;

  const directions = [
    [0, 1],
    [0, -1],
    [1, 0],
    [-1, 0]
  ];

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      for (let [dr, dc] of directions) {
        let match = true;
        for (let i = 0; i < len; i++) {
          const nr = r + dr * i;
          const nc = c + dc * i;
          if (nr < 0 || nr >= rows || nc < 0 || nc >= cols || matrix[nr][nc] !== word[i]) {
            match = false;
            break;
          }
        }
        if (match) {
          return [
            [r, c],
            [r + dr * (len - 1), c + dc * (len - 1)]
          ];
        }
      }
    }
  }
}
```
