---
id: 68f6587287ad1f4ad39b0c7e
title: "チャレンジ 92: 拡張子抽出器"
challengeType: 28
dashedName: challenge-92
---

# --description--

ファイル名を表す文字列が与えられたとき、そのファイルの拡張子を返してください。

- 拡張子とは、ファイル名の最後のピリオド（`.`）の後に続く部分です。
- ファイル名にピリオドが含まれていないか、ピリオドで終わっている場合は `"none"` を返してください。
- 拡張子は大文字・小文字をそのまま保持して返してください。

# --hints--

`getExtension("document.txt")` は `"txt"` を返すべきです。

```js
assert.equal(getExtension("document.txt"), "txt");
```

`getExtension("README")` は `"none"` を返すべきです。

```js
assert.equal(getExtension("README"), "none");
```

`getExtension("image.PNG")` は `"PNG"` を返すべきです。

```js
assert.equal(getExtension("image.PNG"), "PNG");
```

`getExtension(".gitignore")` は `"gitignore"` を返すべきです。

```js
assert.equal(getExtension(".gitignore"), "gitignore");
```

`getExtension("archive.tar.gz")` は `"gz"` を返すべきです。

```js
assert.equal(getExtension("archive.tar.gz"), "gz");
```

`getExtension("final.draft.")` は `"none"` を返すべきです。

```js
assert.equal(getExtension("final.draft."), "none");
```

# --seed--

## --seed-contents--

```js
function getExtension(filename) {

  return filename;
}
```

# --solutions--

```js
function getExtension(filename) {
  const lastDot = filename.lastIndexOf('.');
  
  if (lastDot === -1 || lastDot === filename.length - 1) {
    return "none";
  }
  
  return filename.slice(lastDot + 1);
}
```
