---
id: 68f6587287ad1f4ad39b0c82
title: "チャレンジ 96: 週末ですか？"
challengeType: 28
dashedName: challenge-96
---

# --description--

`"YYYY-MM-DD"` 形式の日付が与えられたとき、週末までの日数を返してください。

- 週末は土曜日から始まります。
- 与えられた日付が土曜日または日曜日の場合は、`"It's the weekend!"` を返してください。
- それ以外の場合は、土曜日までの日数である `"X days until the weekend."` を使って `X` を返してください。
- `X` が `1` の場合は、複数形の `"day"` ではなく単数形の `"days"` を使ってください。
- 計算はローカルタイムゾーンを無視するようにしてください。

# --hints--

`daysUntilWeekend("2025-11-14")` は `"1 day until the weekend."` を返すべきです。

```js
assert.equal(daysUntilWeekend("2025-11-14"), "1 day until the weekend.");
```

`daysUntilWeekend("2025-01-01")` は `"3 days until the weekend."` を返すべきです。

```js
assert.equal(daysUntilWeekend("2025-01-01"), "3 days until the weekend.");
```

`daysUntilWeekend("2025-12-06")` は `"It's the weekend!"` を返すべきです。

```js
assert.equal(daysUntilWeekend("2025-12-06"), "It's the weekend!");
```

`daysUntilWeekend("2026-01-27")` は `"4 days until the weekend."` を返すべきです。

```js
assert.equal(daysUntilWeekend("2026-01-27"), "4 days until the weekend.");
```

`daysUntilWeekend("2026-09-07")` は `"5 days until the weekend."` を返すべきです。

```js
assert.equal(daysUntilWeekend("2026-09-07"), "5 days until the weekend.");
```

`daysUntilWeekend("2026-11-29")` は `"It's the weekend!"` を返すべきです。

```js
assert.equal(daysUntilWeekend("2026-11-29"), "It's the weekend!");
```

# --seed--

## --seed-contents--

```js
function daysUntilWeekend(dateString) {

  return dateString;
}
```

# --solutions--

```js
function daysUntilWeekend(dateString) {
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  const dayOfWeek = date.getUTCDay();

  if (dayOfWeek === 6 || dayOfWeek === 0) {
    return "It's the weekend!";
  }

  const daysUntilSaturday = (6 - dayOfWeek);
  return `${daysUntilSaturday} day${daysUntilSaturday === 1 ? "" : "s"} until the weekend.`;
}
```
