---
id: 68ffb91507a5b645769328ca
title: "チャレンジ 107: FizzBuzz"
challengeType: 28
dashedName: challenge-107
---

# --description--

整数 (`n`) が与えられたとき、`1` から `n` までの整数の配列を返してください（両端を含む）。ただし、次の条件に当てはまる数字はそれぞれ置き換えます：

- 3の倍数は `"Fizz"` に置き換えます。
- 5の倍数は `"Buzz"` に置き換えます。
- 3と5の両方の倍数は `"FizzBuzz"` に置き換えます。

# --hints--

`fizzBuzz(2)` は `[1, 2]` を返すべきです。

```js
assert.deepEqual(fizzBuzz(2), [1, 2]);
```

`fizzBuzz(4)` は `[1, 2, "Fizz", 4]` を返すべきです。

```js
assert.deepEqual(fizzBuzz(4), [1, 2, "Fizz", 4]);
```

`fizzBuzz(8)` は `[1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8]` を返すべきです。

```js
assert.deepEqual(fizzBuzz(8), [1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8]);
```

`fizzBuzz(20)` は `[1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz"]` を返すべきです。

```js
assert.deepEqual(fizzBuzz(20), [1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz"]);
```

`fizzBuzz(50)` は `[1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"]` を返すべきです。

```js
assert.deepEqual(fizzBuzz(50), [1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"]);
```

# --seed--

## --seed-contents--

```js
function fizzBuzz(n) {

  return n;
}
```

# --solutions--

```js
function fizzBuzz(n) {
  const result = [];
  for (let i = 1; i <= n; i++) {
    if (i % 15 === 0) {
      result.push("FizzBuzz");
    } else if (i % 3 === 0) {
      result.push("Fizz");
    } else if (i % 5 === 0) {
      result.push("Buzz");
    } else {
      result.push(i);
    }
  }
  return result;
}
```
