---
id: 68ffb91507a5b645769328c4
title: "チャレンジ101：マークダウン見出しコンバーター"
challengeType: 28
dashedName: challenge-101
---

# --description--

マークダウンの見出しを表す文字列が与えられたら、それと同じ意味のHTML見出しを返してください。

有効なマークダウン見出しは以下の条件を満たす必要があります：

- 先頭に0個以上のスペースがあり、その後に
- 1から6個の連続したハッシュ記号（`#`）が続き、
- その後に少なくとも1つのスペースがあり、最後に
- 見出しのテキストが続く。

ハッシュ記号の数が見出しレベルを決めます。例えば、ハッシュ記号1つは`h1`タグに対応し、6つは`h6`タグに対応します。

もし与えられた文字列が上記の正確な形式でなければ、`"Invalid format"`を返してください。

例えば、`"# My level 1 heading"`が与えられたら、`"<h1>My level 1 heading</h1>"`を返します。

注意：コンソールは文字列内のHTMLタグを表示しない場合があります。タグを含むログを確認するにはブラウザのコンソールをチェックしてください。

# --hints--

`convert("# My level 1 heading")`は`"<h1>My level 1 heading</h1>"`を返すべきです。

```js
assert.equal(convert("# My level 1 heading"), "<h1>My level 1 heading</h1>");
```

`convert("My heading")`は`"Invalid format"`を返すべきです。

```js
assert.equal(convert("My heading"), "Invalid format");
```

`convert("##### My level 5 heading")`は`"<h5>My level 5 heading</h5>"`を返すべきです。

```js
assert.equal(convert("##### My level 5 heading"), "<h5>My level 5 heading</h5>");
```

`convert("#My heading")`は`"Invalid format"`を返すべきです。

```js
assert.equal(convert("#My heading"), "Invalid format");
```

`convert("  ###  My level 3 heading")`は`"<h3>My level 3 heading</h3>"`を返すべきです。

```js
assert.equal(convert("  ###  My level 3 heading"), "<h3>My level 3 heading</h3>");
```

`convert("####### My level 7 heading")`は`"Invalid format"`を返すべきです。

```js
assert.equal(convert("####### My level 7 heading"), "Invalid format");
```

`convert("## My #2 heading")`は`"<h2>My #2 heading</h2>"`を返すべきです。

```js
assert.equal(convert("## My #2 heading"), "<h2>My #2 heading</h2>");
```

# --seed--

## --seed-contents--

```js
function convert(heading) {

  return heading;
}
```

# --solutions--

```js
function convert(heading) {
  const trimmed = heading.trim()

  if (!/^#{1,6}\s+\S+/.test(trimmed)) return "Invalid format"

  const hashCount = trimmed.indexOf(' ');
  const text = trimmed.slice(hashCount).trimStart();

  return `<h${hashCount}>${text}</h${hashCount}>`;
}
```
