---
id: 68ffb91507a5b645769328cc
title: "チャレンジ109: 私の年齢はいくつ？"
challengeType: 28
dashedName: challenge-109
---

# --description--

`YYYY-MM-DD` の形式で与えられた誕生日の日付から、2025年11月27日時点での年齢を返してください。

- すべての誕生日は2025年11月27日より前の有効な日付と仮定してください。
- 年齢は整数で返してください。
- 2025年にすでに誕生日を迎えたかどうかを必ず考慮してください。

# --hints--

`calculateAge("2000-11-20")` は `25` を返すべきです。

```js
assert.equal(calculateAge("2000-11-20"), 25);
```

`calculateAge("2000-12-01")` は `24` を返すべきです。

```js
assert.equal(calculateAge("2000-12-01"), 24);
```

`calculateAge("2014-10-25")` は `11` を返すべきです。

```js
assert.equal(calculateAge("2014-10-25"), 11);
```

`calculateAge("1994-01-06")` は `31` を返すべきです。

```js
assert.equal(calculateAge("1994-01-06"), 31);
```

`calculateAge("1994-12-14")` は `30` を返すべきです。

```js
assert.equal(calculateAge("1994-12-14"), 30);
```

# --seed--

## --seed-contents--

```js
function calculateAge(birthday) {

  return birthday;
}
```

# --solutions--

```js
function calculateAge(birthday) {
  const today = new Date("2025-11-27");
  const [year, month, day] = birthday.split("-").map(Number);
  const birthDate = new Date(year, month - 1, day);

  let age = today.getFullYear() - birthDate.getFullYear();

  if (
    today.getMonth() < birthDate.getMonth() ||
    (today.getMonth() === birthDate.getMonth() && today.getDate() < birthDate.getDate())
  ) {
    age--;
  }

  return age;
}
```
