---
id: 68ffb91507a5b645769328c6
title: "チャレンジ 103: 最小公倍数"
challengeType: 28
dashedName: challenge-103
---

# --description--

2つの整数が与えられたとき、その2つの数の最小公倍数（LCM）を返してください。

2つの数の最小公倍数とは、両方の数の倍数である最小の正の整数です。例えば、`4` と `6` が与えられた場合、`12` を返します。なぜなら：

- `4` の倍数は `4`、`8`、`12` などです。
- `6` の倍数は `6`、`12`、`18` などです。
- `12` は両方の数の倍数である最小の数です。

# --hints--

`lcm(4, 6)` は `12` を返すべきです。

```js
assert.equal(lcm(4, 6), 12);
```

`lcm(9, 6)` は `18` を返すべきです。

```js
assert.equal(lcm(9, 6), 18);
```

`lcm(10, 100)` は `100` を返すべきです。

```js
assert.equal(lcm(10, 100), 100);
```

`lcm(13, 17)` は `221` を返すべきです。

```js
assert.equal(lcm(13, 17), 221);
```

`lcm(45, 70)` は `630` を返すべきです。

```js
assert.equal(lcm(45, 70), 630);
```

# --seed--

## --seed-contents--

```js
function lcm(a, b) {

  return a;
}
```

# --solutions--

```js
function lcm(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  return Math.abs(a * b) / gcd(a, b);
}
```
