---
id: 68ffb91507a5b645769328c5
title: "チャレンジ 102: 最長の単語"
challengeType: 28
dashedName: challenge-102
---

# --description--

文章の文字列が与えられたら、その文章の中で最も長い単語を返してください。

- 単語はスペース1つで区切られています。
- 単語の長さにカウントされるのは文字（`a-z`、大文字小文字は区別しません）のみです。
- 同じ長さの単語が複数ある場合は、最初に出てくる単語を返してください。
- 単語は与えられた文字列に現れる通りに返し、句読点は取り除いてください。

# --hints--

`longestWord("The quick red fox")` は `"quick"` を返すべきです。

```js
assert.equal(longestWord("The quick red fox"), "quick");
```

`longestWord("Hello coding challenge.")` は `"challenge"` を返すべきです。

```js
assert.equal(longestWord("Hello coding challenge."), "challenge");
```

`longestWord("Do Try This At Home.")` は `"This"` を返すべきです。

```js
assert.equal(longestWord("Do Try This At Home."), "This");
```

`longestWord("This sentence... has commas, ellipses, and an exclamation point!")` は `"exclamation"` を返すべきです。

```js
assert.equal(longestWord("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation");
```

`longestWord("A tie? No way!")` は `"tie"` を返すべきです。

```js
assert.equal(longestWord("A tie? No way!"), "tie");
```

`longestWord("Wouldn't you like to know.")` は `"Wouldnt"` を返すべきです。

```js
assert.equal(longestWord("Wouldn't you like to know."), "Wouldnt");
```

# --seed--

## --seed-contents--

```js
function longestWord(sentence) {

  return sentence;
}
```

# --solutions--

```js
function longestWord(sentence) {
  const words = sentence.split(" ");
  let longest = "";

  for (let word of words) {
    const cleaned = word.replace(/[^a-z]/gi, "");
    if (cleaned.length > longest.length) {
      longest = cleaned;
    }
  }

  return longest;
}
```
