---
id: 68ffb91507a5b645769328c7
title: "チャレンジ 104: レシピの分量調整"
challengeType: 28
dashedName: challenge-104
---

# --description--

レシピの材料の配列と分量を調整する数値が与えられたとき、その数値に応じて分量を調整した配列を返してください。

- 与えられる配列の各要素は `"quantity unit ingredient"` の形式の文字列です。例えば `"2 C Flour"` のようになります。
- 分量は与えられた数値でスケールしてください。末尾のゼロは含めず、単位は変換しないでください。
- スケールしたアイテムは元の順番のまま返してください。

# --hints--

`scaleRecipe(["2 C Flour", "1.5 T Sugar"], 2)` は `["4 C Flour", "3 T Sugar"]` を返すべきです。

```js
assert.deepEqual(scaleRecipe(["2 C Flour", "1.5 T Sugar"], 2), ["4 C Flour", "3 T Sugar"]);
```

`scaleRecipe(["4 T Flour", "1 C Milk", "2 T Oil"], 1.5)` は `["6 T Flour", "1.5 C Milk", "3 T Oil"]` を返すべきです。

```js
assert.deepEqual(scaleRecipe(["4 T Flour", "1 C Milk", "2 T Oil"], 1.5), ["6 T Flour", "1.5 C Milk", "3 T Oil"]);
```

`scaleRecipe(["3 C Milk", "2 C Oats"], 0.5)` は `["1.5 C Milk", "1 C Oats"]` を返すべきです。

```js
assert.deepEqual(scaleRecipe(["3 C Milk", "2 C Oats"], 0.5), ["1.5 C Milk", "1 C Oats"]);
```

`scaleRecipe(["2 C All-purpose Flour", "1 t Baking Soda", "1 t Salt", "1 C Butter", "0.5 C Sugar", "0.5 C Brown Sugar", "1 t Vanilla Extract", "2 C Chocolate Chips"], 2.5)` は `["5 C All-purpose Flour", "2.5 t Baking Soda", "2.5 t Salt", "2.5 C Butter", "1.25 C Sugar", "1.25 C Brown Sugar", "2.5 t Vanilla Extract", "5 C Chocolate Chips"]` を返すべきです。

```js
assert.deepEqual(scaleRecipe(["2 C All-purpose Flour", "1 t Baking Soda", "1 t Salt", "1 C Butter", "0.5 C Sugar", "0.5 C Brown Sugar", "1 t Vanilla Extract", "2 C Chocolate Chips"], 2.5), ["5 C All-purpose Flour", "2.5 t Baking Soda", "2.5 t Salt", "2.5 C Butter", "1.25 C Sugar", "1.25 C Brown Sugar", "2.5 t Vanilla Extract", "5 C Chocolate Chips"]);
```

# --seed--

## --seed-contents--

```js
function scaleRecipe(ingredients, scale) {

  return ingredients;
}
```

# --solutions--

```js
function scaleRecipe(ingredients, scale) {
  return ingredients.map(item => {
    const [quantityStr, unit, ...rest] = item.split(" ");
    const ingredientName = rest.join(" ");
    
    const quantity = parseFloat(quantityStr) * scale;
    
    return `${quantity} ${unit} ${ingredientName}`;
  });
}
```
