---
id: 68ffb91507a5b645769328cb
title: "チャレンジ 108: BuzzFizz"
challengeType: 28
dashedName: challenge-108
---

# --description--

配列が与えられたとき、その配列が1から最後の要素までの正しいFizzBuzzシーケンスかどうかを判定してください。シーケンスが正しいとは以下の条件を満たすことです：

- 3の倍数の数字は "Fizz" に置き換えられている
- 5の倍数の数字は "Buzz" に置き換えられている
- 3と5の両方の倍数の数字は "FizzBuzz" に置き換えられている
- その他の数字は1から始まる昇順の整数のままである
- 配列は1から始まり、欠けている要素や余分な要素がない

# --hints--

`isFizzBuzz([1, 2, "Fizz", 4])` は `true` を返すべきです。

```js
assert.isTrue(isFizzBuzz([1, 2, "Fizz", 4]));
```

`isFizzBuzz([1, 2, 3, 4])` は `false` を返すべきです。

```js
assert.isFalse(isFizzBuzz([1, 2, 3, 4]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", 7])` は `false` を返すべきです。

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", 7]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "FizzBuzz"])` は `false` を返すべきです。

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "FizzBuzz"]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Fizz"])` は `false` を返すべきです。

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Fizz"]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Buzz"])` は `false` を返すべきです。

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Buzz"]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"])` は `true` を返すべきです。

```js
assert.isTrue(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"]));
```

# --seed--

## --seed-contents--

```js
function isFizzBuzz(sequence) {

  return sequence;
}
```

# --solutions--

```js
function isFizzBuzz(sequence) {
  for (let i = 0; i < sequence.length; i++) {
    const n = i + 1;
    if (n % 15 === 0 && sequence[i] !== "FizzBuzz") return false;
    if (n % 3 === 0 && n % 5 !== 0 && sequence[i] !== "Fizz") return false;
    if (n % 5 === 0 && n % 3 !== 0 && sequence[i] !== "Buzz") return false;
    if (n % 3 !== 0 && n % 5 !== 0 && sequence[i] !== n) return false;
  }
  return true;
}
```
