---
id: 68ffb91507a5b645769328c3
title: "チャレンジ100：100文字"
challengeType: 28
dashedName: challenge-100
---

# --description--

100回目のデイリーコーディングチャレンジへようこそ！

文字列が与えられたら、その文字を繰り返して結果がちょうど100文字になるようにしてください。繰り返しで100文字を超えた場合は、余分な部分を切り詰めてちょうど100文字にしてください。

# --hints--

`oneHundred("One hundred ")` は `"One hundred One hundred One hundred One hundred One hundred One hundred One hundred One hundred One "` を返すべきです。

```js
assert.equal(oneHundred("One hundred "), "One hundred One hundred One hundred One hundred One hundred One hundred One hundred One hundred One ");
```

`oneHundred("freeCodeCamp ")` は `"freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeC"` を返すべきです。

```js
assert.equal(oneHundred("freeCodeCamp "), "freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeC");
```

`oneHundred("daily challenges ")` は `"daily challenges daily challenges daily challenges daily challenges daily challenges daily challenge"` を返すべきです。

```js
assert.equal(oneHundred("daily challenges "), "daily challenges daily challenges daily challenges daily challenges daily challenges daily challenge");
```

`oneHundred("!")` は `"!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!"` を返すべきです。

```js
assert.equal(oneHundred("!"), "!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!");
```

# --seed--

## --seed-contents--

```js
function oneHundred(chars) {

  return chars;
}
```

# --solutions--

```js
function oneHundred(chars) {
  let result = "";
  let i = 0;

  while (result.length < 100) {
    result += chars[i % chars.length];
    i++;
  }

  return result.slice(0, 100);
}
```
