---
id: 68ffb91507a5b645769328c9
title: "チャレンジ 106: メッセージバリデーター"
challengeType: 28
dashedName: challenge-106
---

# --description--

メッセージの文字列とバリデーションの文字列が与えられたとき、メッセージが有効かどうかを判定してください。

- メッセージが有効なのは、メッセージ内の各単語がバリデーション文字列の対応する文字で始まっている場合です。順番も一致している必要があります。
- 文字の大文字・小文字は区別しません。
- メッセージ内の単語は単一のスペースで区切られています。

# --hints--

`isValidMessage("hello world", "hw")` は `true` を返すべきです。

```js
assert.isTrue(isValidMessage("hello world", "hw"));
```

`isValidMessage("ALL CAPITAL LETTERS", "acl")` は `true` を返すべきです。

```js
assert.isTrue(isValidMessage("ALL CAPITAL LETTERS", "acl"));
```

`isValidMessage("Coding challenge are boring.", "cca")` は `false` を返すべきです。

```js
assert.isFalse(isValidMessage("Coding challenge are boring.", "cca"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD")` は `true` を返すべきです。

```js
assert.isTrue(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT")` は `false` を返すべきです。

```js
assert.isFalse(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT"));
```

# --seed--

## --seed-contents--

```js
function isValidMessage(message, validator) {

  return message;
}
```

# --solutions--

```js
function isValidMessage(message, validation) {
  const words = message.split(" ");
  
  if (words.length !== validation.length) return false;

  for (let i = 0; i < words.length; i++) {
    if (words[i][0].toLowerCase() !== validation[i].toLowerCase()) {
      return false;
    }
  }

  return true;
}
```
