---
id: 68ffb91507a5b645769328c8
title: "チャレンジ 105: 文字数カウント"
challengeType: 28
dashedName: challenge-105
---

# --description--

与えられた文章の文字列から、アルファベット順に各文字の出現数を配列で返してください。

- 大文字と小文字は同じ文字として扱います。
- 数字、空白、句読点などは無視してください。
- カウントと文字は `"letter count"` の形式で返してください。例えば、`"a 3"` のようにします。
- 返す文字はすべて小文字にしてください。
- 文字列に含まれていない文字のカウントは返さないでください。

# --hints--

`countCharacters("hello world")` は `["d 1", "e 1", "h 1", "l 3", "o 2", "r 1", "w 1"]` を返すべきです。

```js
assert.deepEqual(countCharacters("hello world"), ["d 1", "e 1", "h 1", "l 3", "o 2", "r 1", "w 1"]);
```

`countCharacters("I love coding challenges!")` は `["a 1", "c 2", "d 1", "e 3", "g 2", "h 1", "i 2", "l 3", "n 2", "o 2", "s 1", "v 1"]` を返すべきです。

```js
assert.deepEqual(countCharacters("I love coding challenges!"), ["a 1", "c 2", "d 1", "e 3", "g 2", "h 1", "i 2", "l 3", "n 2", "o 2", "s 1", "v 1"]);
```

`countCharacters("// TODO: Complete this challenge ASAP!")` は `["a 3", "c 2", "d 1", "e 4", "g 1", "h 2", "i 1", "l 3", "m 1", "n 1", "o 3", "p 2", "s 2", "t 3"]` を返すべきです。

```js
assert.deepEqual(countCharacters("// TODO: Complete this challenge ASAP!"), ["a 3", "c 2", "d 1", "e 4", "g 1", "h 2", "i 1", "l 3", "m 1", "n 1", "o 3", "p 2", "s 2", "t 3"]);
```

# --seed--

## --seed-contents--

```js
function countCharacters(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countCharacters(sentence) {
  const counts = {};

  for (let char of sentence.toLowerCase()) {
    if (char >= "a" && char <= "z") {
      counts[char] = (counts[char] || 0) + 1;
    }
  }

  return Object.keys(counts)
               .sort()
               .map(letter => `${letter} ${counts[letter]}`);
}
```
