---
id: 69162d64f96574d9bb629f00
title: "チャレンジ 115: Markdown 番号付きリスト項目コンバーター"
challengeType: 28
dashedName: challenge-115
---

# --description--

Markdownで表された番号付きリストの項目を表す文字列が与えられたとき、それに対応するHTMLの文字列を返してください。

Markdownの有効な番号付きリスト項目は以下の条件を満たす必要があります：

- 0個以上のスペースで始まり、
- 1以上の数字とピリオド（`.`）が続き、
- その後に少なくとも1つのスペースがあり、
- リスト項目のテキストが続く。

文字列が上記の正確な形式でない場合は`"Invalid format"`を返してください。そうでなければ、リスト項目のテキストを`li`タグで囲んだ文字列を返してください。

例えば、`"1. My item"`が与えられたら、`"<li>My item</li>"`を返します。

注意：コンソールは文字列内のHTMLタグをログに表示しない場合があります。タグを含むログを見るにはブラウザのコンソールを確認してください。

# --hints--

`convertListItem("1. My item")`は`"<li>My item</li>"`を返すべきです。

```js
assert.equal(convertListItem("1. My item"), "<li>My item</li>");
```

`convertListItem(" 1.  Another item")`は`"<li>Another item</li>"`を返すべきです。

```js
assert.equal(convertListItem(" 1.  Another item"), "<li>Another item</li>");
```

`convertListItem("1 . invalid item")`は`"Invalid format"`を返すべきです。

```js
assert.equal(convertListItem("1 . invalid item"), "Invalid format");
```

`convertListItem("2. list item text")`は`"<li>list item text</li>"`を返すべきです。

```js
assert.equal(convertListItem("2. list item text"), "<li>list item text</li>");
```

`convertListItem(". invalid again")`は`"Invalid format"`を返すべきです。

```js
assert.equal(convertListItem(". invalid again"), "Invalid format");
```

`convertListItem("A. last invalid")`は`"Invalid format"`を返すべきです。

```js
assert.equal(convertListItem("A. last invalid"), "Invalid format");
```

# --seed--

## --seed-contents--

```js
function convertListItem(markdown) {

  return markdown;
}
```

# --solutions--

```js
function convertListItem(markdown) {
  const match = markdown.match(/^\s*(\d+)\.\s+(.+)$/);
  if (!match) return "Invalid format";

  const text = match[2];
  return `<li>${text}</li>`;
}
```
