---
id: 69162d64f96574d9bb629efd
title: "チャレンジ 112: AI検出器"
challengeType: 28
dashedName: challenge-112
---

# --description--

本日のチャレンジは、2022年11月30日にリリースされたChatGPTに触発されています。

1つ以上の文からなる文字列が与えられたとき、以下のルールに基づいてAIによって生成された可能性が高いかどうかを判定してください。

- 2つ以上のダッシュ（`-`）が含まれている。
- 2つ以上の括弧のセット（`()`）が含まれている。括弧内にテキストがあってもよい。
- 7文字以上の単語が3つ以上含まれている。

- 単語は1つのスペースで区切られ、文字（`A-Z`）のみで構成されます。句読点やその他の文字は単語の一部として含めないでください。

与えられた文が上記のいずれかのルールを満たす場合は`"AI"`を返し、そうでなければ`"Human"`を返してください。

# --hints--

`detectAI("The quick brown fox jumped over the lazy dog.")`は`"Human"`を返すべきです。

```js
assert.equal(detectAI("The quick brown fox jumped over the lazy dog."), "Human");
```

`detectAI("The hypersonic brown fox - jumped (over) the lazy dog.")`は`"Human"`を返すべきです。

```js
assert.equal(detectAI("The hypersonic brown fox - jumped (over) the lazy dog."), "Human");
```

`detectAI("Yes - you're right! I made a mistake there - let me try again.")`は`"AI"`を返すべきです。

```js
assert.equal(detectAI("Yes - you're right! I made a mistake there - let me try again."), "AI");
```

`detectAI("The extraordinary students were studying vivaciously.")`は`"AI"`を返すべきです。

```js
assert.equal(detectAI("The extraordinary students were studying vivaciously."), "AI");
```

`detectAI("The (excited) student was (coding) in the library.")`は`"AI"`を返すべきです。

```js
assert.equal(detectAI("The (excited) student was (coding) in the library."), "AI");
```

# --seed--

## --seed-contents--

```js
function detectAI(text) {

  return text;
}
```

# --solutions--

```js
function detectAI(text) {
  const dashCount = (text.match(/-/g) || []).length;
  if (dashCount >= 2) return "AI";

  const parenCount = (text.match(/\([^)]*\)/g) || []).length;
  if (parenCount >= 2) return "AI";

  const words = text.split(" ");
  const longWordCount = words.filter(w => w.replace(/[^A-Za-z]/g, "").length >= 7).length;
  if (longWordCount >= 3) return "AI";

  return "Human";
}
```
