---
id: 69162d64f96574d9bb629f01
title: "チャレンジ 116: 順列の数"
challengeType: 28
dashedName: challenge-116
---

# --description--

文字列が与えられたとき、その文字から作成できる異なる順列の数を返してください。

- 順列とは、文字列の文字を並べ替えたものです。
- 重複する順列は数えないでください。
- 文字列に繰り返しの文字が含まれている場合、同じ並びは一度だけ数えます。
- 文字列は文字（`A-Z`、`a-z`）のみを含みます。

たとえば、`"abb"`が与えられた場合、文字の並べ方が3通りあるので`3`を返してください：`"abb"`、`"bab"`、および`"bba"`。

# --hints--

`countPermutations("abb")`は`3`を返すべきです。

```js
assert.equal(countPermutations("abb"), 3);
```

`countPermutations("abc")`は`6`を返すべきです。

```js
assert.equal(countPermutations("abc"), 6);
```

`countPermutations("racecar")`は`630`を返すべきです。

```js
assert.equal(countPermutations("racecar"), 630);
```

`countPermutations("freecodecamp")`は`39916800`を返すべきです。

```js
assert.equal(countPermutations("freecodecamp"), 39916800);
```

# --seed--

## --seed-contents--

```js
function countPermutations(str) {

  return str;
}
```

# --solutions--

```js
function countPermutations(str) {
  const freq = {};
  for (const ch of str) {
    freq[ch] = (freq[ch] || 0) + 1;
  }

  function factorial(n) {
    let result = 1;
    for (let i = 2; i <= n; i++) result *= i;
    return result;
  }

  const n = str.length;
  let result = factorial(n);

  for (const ch in freq) {
    result /= factorial(freq[ch]);
  }

  return result;
}
```
