---
id: 69162d64f96574d9bb629f04
title: "チャレンジ 119: 文字列圧縮"
challengeType: 28
dashedName: challenge-119
---

# --description--

文字列 sentence が与えられたとき、連続して重複する単語を、その単語に続けて丸括弧内に繰り返し回数を付けた形に置き換えた圧縮版の文章を返してください。

- 圧縮されるのは連続した重複のみです。
- 単語は単一のスペースで区切られています。

例えば、`"yes yes yes please"` が与えられたら、`"yes(3) please"` を返してください。

# --hints--

`compressString("yes yes yes please")` は `"yes(3) please"` を返すべきです。

```js
assert.equal(compressString("yes yes yes please"), "yes(3) please");
```

`compressString("I have have have apples")` は `"I have(3) apples"` を返すべきです。

```js
assert.equal(compressString("I have have have apples"), "I have(3) apples");
```

`compressString("one one three and to the the the the")` は `"one(2) three and to the(4)"` を返すべきです。

```js
assert.equal(compressString("one one three and to the the the the"), "one(2) three and to the(4)");
```

`compressString("route route route route route route tee tee tee tee tee tee")` は `"route(6) tee(6)"` を返すべきです。

```js
assert.equal(compressString("route route route route route route tee tee tee tee tee tee"), "route(6) tee(6)");
```

# --seed--

## --seed-contents--

```js
function compressString(sentence) {

  return sentence;
}
```

# --solutions--

```js
function compressString(sentence) {
  const words = sentence.split(" ");
  const result = [];
  let count = 1;

  for (let i = 0; i < words.length; i++) {
    if (words[i] === words[i + 1]) {
      count++;
    } else {
      if (count > 1) {
        result.push(`${words[i]}(${count})`);
      } else {
        result.push(words[i]);
      }
      count = 1;
    }
  }

  return result.join(" ");
}
```
