---
id: 69162d64f96574d9bb629f02
title: "チャレンジ117：対称差"
challengeType: 28
dashedName: challenge-117
---

# --description--

2つの配列が与えられたとき、それらの対称差を含む新しい配列を返してください。

- 2つの集合の対称差とは、どちらか一方の集合にだけ現れる値の集合です。
- 入力配列に最初に現れた順番で値を返してください。

# --hints--

`difference([1, 2, 3], [3, 4, 5])` は `[1, 2, 4, 5]` を返すべきです。

```js
assert.deepEqual(difference([1, 2, 3], [3, 4, 5]), [1, 2, 4, 5]);
```

`difference(["a", "b"], ["c", "b"])` は `["a", "c"]` を返すべきです。

```js
assert.deepEqual(difference(["a", "b"], ["c", "b"]), ["a", "c"]);
```

`difference([1, "a", 2], [2, "b", "a"])` は `[1, "b"]` を返すべきです。

```js
assert.deepEqual(difference([1, "a", 2], [2, "b", "a"]), [1, "b"]);
```

`difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9])` は `[2, 4, 6, 8]` を返すべきです。

```js
assert.deepEqual(difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9]), [2, 4, 6, 8]);
```

# --seed--

## --seed-contents--

```js
function difference(arr1, arr2) {

  return arr1;
}
```

# --solutions--

```js
function difference(arr1, arr2) {
  const diff = [];

  for (const v of arr1) {

    if (!diff.includes(v) && !arr2.includes(v)) diff.push(v)
  }
  
  for (const v of arr2) {
    if (!diff.includes(v) && !arr1.includes(v)) diff.push(v)
  }

  return diff;
}
```
