---
id: 69162d64f96574d9bb629efe
title: "チャレンジ 113: マイルからキロメートルへ"
challengeType: 28
dashedName: challenge-113
---

# --description--

マイルで表された距離が数値で与えられたとき、それに相当するキロメートルの距離を返してください。

- 入力は常に0以上の数値です。
- 1マイルは1.60934キロメートルに相当します。
- 結果は小数点以下2桁に丸めてください。
- 丸めた結果の不要な末尾のゼロは削除してください。

# --hints--

`convertToKm(1)` は `1.61` を返すべきです。

```js
assert.equal(convertToKm(1), 1.61);
```

`convertToKm(21)` は `33.8` を返すべきです。

```js
assert.equal(convertToKm(21), 33.8);
```

`convertToKm(3.5)` は `5.63` を返すべきです。

```js
assert.equal(convertToKm(3.5), 5.63);
```

`convertToKm(0)` は `0` を返すべきです。

```js
assert.equal(convertToKm(0), 0);
```

`convertToKm(0.621371)` は `1` を返すべきです。

```js
assert.equal(convertToKm(0.621371), 1);
```

# --seed--

## --seed-contents--

```js
function convertToKm(miles) {

  return miles;
}
```

# --solutions--

```js
function convertToKm(miles) {
  const km = miles * 1.60934;
  return Math.round(km * 100) / 100;
}
```
