---
id: 69162d64f96574d9bb629f03
title: "チャレンジ 118: 日付フォーマッター"
challengeType: 28
dashedName: challenge-118
---

# --description--

`"Month day, year"` の形式で与えられた日付を、`"YYYY-MM-DD"` の形式で返してください。

- 与えられる月は英語の完全な月名です。例えば、`"January"`、`"February"`などです。
- 戻り値では、月と日が2桁になるように必要に応じて先頭にゼロを付けてください。

例えば、`"December 6, 2025"` が与えられたら、`"2025-12-06"` を返してください。

# --hints--

`formatDate("December 6, 2025")` は `"2025-12-06"` を返すべきです。

```js
assert.equal(formatDate("December 6, 2025"), "2025-12-06");
```

`formatDate("January 1, 2000")` は `"2000-01-01"` を返すべきです。

```js
assert.equal(formatDate("January 1, 2000"), "2000-01-01");
```

`formatDate("November 11, 1111")` は `"1111-11-11"` を返すべきです。

```js
assert.equal(formatDate("November 11, 1111"), "1111-11-11");
```

`formatDate("September 7, 512")` は `"512-09-07"` を返すべきです。

```js
assert.equal(formatDate("September 7, 512"), "512-09-07");
```

`formatDate("May 4, 1950")` は `"1950-05-04"` を返すべきです。

```js
assert.equal(formatDate("May 4, 1950"), "1950-05-04");
```

`formatDate("February 29, 1992")` は `"1992-02-29"` を返すべきです。

```js
assert.equal(formatDate("February 29, 1992"), "1992-02-29");
```

# --seed--

## --seed-contents--

```js
function formatDate(dateString) {

  return dateString;
}
```

# --solutions--

```js
function formatDate(dateString) {
  const months = {
    January: "01",
    February: "02",
    March: "03",
    April: "04",
    May: "05",
    June: "06",
    July: "07",
    August: "08",
    September: "09",
    October: "10",
    November: "11",
    December: "12"
  };

  const [month, day, year] = dateString.replace(",", "").split(" ");

  return `${year}-${months[month]}-${day.padStart(2, "0")}`;
}
```
