---
id: 69162d64f96574d9bb629efb
title: "チャレンジ 110: 単語当てゲーム"
challengeType: 28
dashedName: challenge-110
---

# --description--

同じ長さの2つの文字列、秘密の単語と推測した単語が与えられたとき、以下のルールで推測した単語を秘密の単語と比較してください。

- 秘密の単語と推測は大文字のアルファベット（`"A"`から`"Z"`）のみで構成されます。
- 推測の各文字について、秘密の単語との一致の仕方に応じて数字に置き換えます：
  - `"2"`：文字が秘密の単語にあり、かつ正しい位置にある場合。
  - `"1"`：文字が秘密の単語にあるが、位置が違う場合。
  - `"0"`：文字が秘密の単語にない場合。
- 秘密の単語の各文字は最大で1回だけ使えます。
- 完全一致（`"2"`）は最初に割り当てられ、その後、部分一致（`"1"`）が左から右へ残りの文字に割り当てられます。
- 推測の中で同じ文字が複数回出現する場合、その文字は秘密の単語に出現する回数までしか一致できません。

例えば、秘密の単語が`"APPLE"`で推測が`"POPPA"`の場合、`"10201"`を返します：

最初の`"P"`は正しい位置ではない（`"1"`）、`"O"`は秘密の単語にない（`"0"`）、2番目の`"P"`は正しい位置にある（`"2"`）、3番目の`"P"`はゼロ（`"0"`）で、秘密の単語の`"P"`が2回使われているため、`"A"`は正しい位置ではない（`"1"`）です。

# --hints--

`compare("APPLE", "POPPA")`は`"10201"`を返すべきです。

```js
assert.equal(compare("APPLE", "POPPA"), "10201");
```

`compare("REACT", "TRACE")`は`"11221"`を返すべきです。

```js
assert.equal(compare("REACT", "TRACE"), "11221");
```

`compare("DEBUGS", "PYTHON")`は`"000000"`を返すべきです。

```js
assert.equal(compare("DEBUGS", "PYTHON"), "000000");
```

`compare("JAVASCRIPT", "TYPESCRIPT")`は`"0000222222"`を返すべきです。

```js
assert.equal(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222");
```

`compare("ORANGE", "ROUNDS")`は`"110200"`を返すべきです。

```js
assert.equal(compare("ORANGE", "ROUNDS"), "110200");
```

`compare("WIRELESS", "ETHERNET")`は`"10021000"`を返すべきです。

```js
assert.equal(compare("WIRELESS", "ETHERNET"), "10021000");
```

# --seed--

## --seed-contents--

```js
function compare(word, guess) {

  return word;
}
```

# --solutions--

```js
function compare(word, guess) {
  const n = word.length;
  let result = Array(n).fill("0");
  let used = Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    if (guess[i] === word[i]) {
      result[i] = "2";
      used[i] = true;
    }
  }

  for (let i = 0; i < n; i++) {
    if (result[i] === "0") {
      for (let j = 0; j < n; j++) {
        if (!used[j] && guess[i] === word[j]) {
          result[i] = "1";
          used[j] = true;
          break;
        }
      }
    }
  }

  return result.join("");
}
```
