---
id: 69162d64f96574d9bb629eff
title: "チャレンジ114：キャメルケースからスネークケースへ"
challengeType: 28
dashedName: challenge-114
---

# --description--

キャメルケースの文字列が与えられたら、次のルールに従ってスネークケースのバージョンを返してください。

- 入力の文字列は文字（`A-Z` と `a-z`）のみを含み、必ず小文字で始まります。
- キャメルケースの文字列の大文字はすべて新しい単語の始まりです。
- すべての文字を小文字に変換してください。
- 単語はアンダースコア（`_`）で区切ってください。

# --hints--

`toSnake("helloWorld")` は `"hello_world"` を返すべきです。

```js
assert.equal(toSnake("helloWorld"), "hello_world");
```

`toSnake("myVariableName")` は `"my_variable_name"` を返すべきです。

```js
assert.equal(toSnake("myVariableName"), "my_variable_name");
```

`toSnake("freecodecampDailyChallenges")` は `"freecodecamp_daily_challenges"` を返すべきです。

```js
assert.equal(toSnake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges");
```

# --seed--

## --seed-contents--

```js
function toSnake(camel) {

  return camel;
}
```

# --solutions--

```js
function toSnake(camel) {
  return camel
    .replace(/([A-Z])/g, "_$1")
    .toLowerCase();
}
```
