---
id: 691b5597f0f3e85a588a5c94
title: "チャレンジ 130: チェッカーボード"
challengeType: 28
dashedName: challenge-130
---

# --description--

2つの数字が入った配列が与えられます。最初の数字が行数、2番目の数字が列数です。指定されたサイズの`"X"`と`"O"`の文字で埋められた行列（配列の配列）を返してください。

- 文字はチェッカーボードのように交互に並べます。
- 左上のセルは常に`"X"`でなければなりません。

例えば、`[3, 3]`が与えられた場合、次のように返します。

```json
[
  ["X", "O", "X"],
  ["O", "X", "O"],
  ["X", "O", "X"]
]
```

# --hints--

`createBoard([3, 3])`は`[["X", "O", "X"], ["O", "X", "O"], ["X", "O", "X"]]`を返すべきです。

```js
assert.deepEqual(createBoard([3, 3]), [["X", "O", "X"], ["O", "X", "O"], ["X", "O", "X"]]);
```

`createBoard([6, 1])`は`[["X"], ["O"], ["X"], ["O"], ["X"], ["O"]]`を返すべきです。

```js
assert.deepEqual(createBoard([6, 1]), [["X"], ["O"], ["X"], ["O"], ["X"], ["O"]]);
```

`createBoard([2, 10])`は`[["X", "O", "X", "O", "X", "O", "X", "O", "X", "O"], ["O", "X", "O", "X", "O", "X", "O", "X", "O", "X"]]`を返すべきです。

```js
assert.deepEqual(createBoard([2, 10]), [["X", "O", "X", "O", "X", "O", "X", "O", "X", "O"], ["O", "X", "O", "X", "O", "X", "O", "X", "O", "X"]]);
```

`createBoard([5, 4])`は`[["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"]]`を返すべきです。

```js
assert.deepEqual(createBoard([5, 4]), [["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"]]);
```

# --seed--

## --seed-contents--

```js
function createBoard(dimensions) {
  return dimensions;
}
```

# --solutions--

```js
function createBoard(dimensions) {
  const [rows, cols] = dimensions;
  const board = [];

  for (let r = 0; r < rows; r++) {
    const row = [];
    for (let c = 0; c < cols; c++) {
      row.push((r + c) % 2 === 0 ? "X" : "O");
    }
    board.push(row);
  }

  return board;
}
```
