---
id: 691b559495c5cb5a37b9b488
title: "チャレンジ128：子音の数"
challengeType: 28
dashedName: challenge-128
---

# --description--

文字列と目標の数が与えられたとき、その文字列にちょうど目標の数だけ子音が含まれているかどうかを判定してください。

- 子音とは、大文字・小文字を問わず`"a"`、`"e"`、`"i"`、`"o"`、`"u"`以外のすべてのアルファベット文字です。
- 数字、句読点、空白、その他の文字以外の文字はカウントしないでください。

# --hints--

`hasConsonantCount("helloworld", 7)`は`true`を返すべきです。

```js
assert.isTrue(hasConsonantCount("helloworld", 7));
```

`hasConsonantCount("eieio", 5)`は`false`を返すべきです。

```js
assert.isFalse(hasConsonantCount("eieio", 5));
```

`hasConsonantCount("freeCodeCamp Rocks!", 11)`は`true`を返すべきです。

```js
assert.isTrue(hasConsonantCount("freeCodeCamp Rocks!", 11));
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)`は`false`を返すべきです。

```js
assert.isFalse(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)
);
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)`は`true`を返すべきです。

```js
assert.isTrue(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)
);
```

# --seed--

## --seed-contents--

```js
function hasConsonantCount(text, target) {
  return text;
}
```

# --solutions--

```js
function hasConsonantCount(text, target) {
  const vowels = new Set(["a", "e", "i", "o", "u"]);
  let count = 0;

  for (const char of text.toLowerCase()) {
    if (/[a-z]/.test(char) && !vowels.has(char)) {
      count++;
    }
  }

  return count === target;
}
```
