---
id: 691b559495c5cb5a37b9b489
title: "チャレンジ 129: Markdown ブロッククオートパーサー"
challengeType: 28
dashedName: challenge-129
---

# --description--

Markdown のブロッククオートを含む文字列が与えられたとき、対応する HTML 文字列を返してください。

Markdown のブロッククオートとは、次の条件を満たす任意の行です。

- 0 個以上のスペースで始まる
- 続いて大なり記号 (`>`)
- その後に 1 個以上のスペース
- 最後にブロッククオートのテキストが続く

ブロッククオートのテキストを、開始および終了の HTML `blockquote` タグで囲んで返してください。

例えば、`"> This is a quote"` が与えられたら、`<blockquote>This is a quote</blockquote>` を返します。

注意: コンソールは文字列内の HTML タグをログに表示しない場合があります。タグを含むログを確認するにはブラウザのコンソールをチェックしてください。

# --hints--

`parseBlockquote("> This is a quote")` は `"<blockquote>This is a quote</blockquote>"` を返すべきです。

```js
assert.equal(parseBlockquote("> This is a quote"), "<blockquote>This is a quote</blockquote>");
```

`parseBlockquote(" > This is also a quote")` は `"<blockquote>This is also a quote</blockquote>"` を返すべきです。

```js
assert.equal(parseBlockquote(" > This is also a quote"), "<blockquote>This is also a quote</blockquote>");
```

`parseBlockquote("  >    So  Is  This")` は `"<blockquote>So  Is  This</blockquote>"` を返すべきです。

```js
assert.equal(parseBlockquote("  >    So  Is  This"), "<blockquote>So  Is  This</blockquote>");
```

# --seed--

## --seed-contents--

```js
function parseBlockquote(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBlockquote(markdown) {
  const blockquoteRegex = /^\s*>\s+(.+)$/;
  const match = markdown.match(blockquoteRegex);

  const text = match[1];
  return `<blockquote>${text}</blockquote>`;
}
```
