---
id: 691b559495c5cb5a37b9b481
title: "チャレンジ 121: 最も頻繁に現れる要素"
challengeType: 28
dashedName: challenge-121
---

# --description--

配列の要素が与えられたとき、最も頻繁に現れる要素を返してください。

- 最も頻繁に現れる要素は必ず1つだけです。

# --hints--

`mostFrequent(["a", "b", "a", "c"])` は `"a"` を返すべきです。

```js
assert.equal(mostFrequent(["a", "b", "a", "c"]), "a");
```

`mostFrequent([2, 3, 5, 2, 6, 3, 2, 7, 2, 9])` は `2` を返すべきです。

```js
assert.equal(mostFrequent([2, 3, 5, 2, 6, 3, 2, 7, 2, 9]), 2);
```

`mostFrequent([true, false, "false", "true", false])` は `false` を返すべきです。

```js
assert.isFalse(mostFrequent([true, false, "false", "true", false]));
```

`mostFrequent([40, 20, 70, 30, 10, 40, 10, 50, 40, 60])` は `40` を返すべきです。

```js
assert.equal(mostFrequent([40, 20, 70, 30, 10, 40, 10, 50, 40, 60]), 40);
```

# --seed--

## --seed-contents--

```js
function mostFrequent(arr) {
  return arr;
}
```

# --solutions--

```js
function mostFrequent(arr) {
  const freq = {};
  let maxCount = 0;
  let mostElem = null;

  for (const el of arr) {
    freq[el] = (freq[el] || 0) + 1;
    if (freq[el] > maxCount) {
      maxCount = freq[el];
      mostElem = el;
    }
  }

  return mostElem;
}
```
