---
id: 691b559495c5cb5a37b9b485
title: "チャレンジ125：ライフゲーム"
challengeType: 28
dashedName: challenge-125
---

# --description--

Conwayのライフゲームの現在の状態を表す行列（配列の配列）が与えられたとき、次のルールに従って行列の次の状態を返してください。

- 各セルは`1`（生きている）か`0`（死んでいる）です。
- セルの隣接セルは、上下左右および斜めの最大8つの周囲のセルです。
- 端のセルは8つ未満の隣接セルしか持ちません。

各セルの更新ルール：

- 生きているセルで隣接する生きているセルが2未満の場合は死にます（過疎）。
- 生きているセルで隣接する生きているセルが2つか3つの場合は生き続けます。
- 生きているセルで隣接する生きているセルが3つより多い場合は死にます（過密）。
- 死んでいるセルで隣接する生きているセルがちょうど3つの場合は生き返ります（再生）。

例えば、以下が与えられた場合：

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

返すべきは：

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

各セルは隣接する生きているセルの数に応じて更新されます。例えば、`[0][0]`は死んだまま（生きている隣接セル2つ）、`[0][1]`は生き続ける（生きている隣接セル2つ）、`[0][2]`は死ぬ（生きている隣接セル3つ）などです。

# --hints--

`gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]])`は`[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`を返すべきです。

```js
assert.deepEqual(gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]]);
```

`gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])`は`[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`を返すべきです。

```js
assert.deepEqual(gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]);
```

`gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]])`は`[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`を返すべきです。

```js
assert.deepEqual(gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]]);
```

`gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])`は`[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`を返すべきです。

```js
assert.deepEqual(gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function gameOfLife(grid) {
  return grid;
}
```

# --solutions--

```js
function gameOfLife(grid) {
  const rows = grid.length;
  const cols = grid[0].length;

  const countLiveNeighbors = (r, c) => {
    let count = 0;
    for (let i = r - 1; i <= r + 1; i++) {
      for (let j = c - 1; j <= c + 1; j++) {
        if (
          i >= 0 &&
          i < rows &&
          j >= 0 &&
          j < cols &&
          !(i === r && j === c) &&
          grid[i][j] === 1
        ) {
          count++;
        }
      }
    }
    return count;
  };

  const next = grid.map(row => [...row]);

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const liveNeighbors = countLiveNeighbors(r, c);

      if (grid[r][c] === 1) {
        if (liveNeighbors < 2 || liveNeighbors > 3) next[r][c] = 0;
      } else {
        if (liveNeighbors === 3) next[r][c] = 1;
      }
    }
  }

  return next;
}
```
