---
id: 691b559495c5cb5a37b9b486
title: "チャレンジ126：大文字に変換する"
challengeType: 28
dashedName: challenge-126
---

# --description--

文字列 title が与えられたとき、次のルールに従ってタイトルケースでフォーマットされた新しい文字列を返してください。

- 各単語の最初の文字を大文字にします。
- 各単語のその他の文字はすべて小文字にします。
- 単語は常に1つのスペースで区切られています。

# --hints--

`titleCase("hello world")` は `"Hello World"` を返すべきです。

```js
assert.equal(titleCase("hello world"), "Hello World");
```

`titleCase("the quick brown fox")` は `"The Quick Brown Fox"` を返すべきです。

```js
assert.equal(titleCase("the quick brown fox"), "The Quick Brown Fox");
```

`titleCase("JAVASCRIPT AND PYTHON")` は `"Javascript And Python"` を返すべきです。

```js
assert.equal(titleCase("JAVASCRIPT AND PYTHON"), "Javascript And Python");
```

`titleCase("AvOcAdO tOAst fOr brEAkfAst")` は `"Avocado Toast For Breakfast"` を返すべきです。

```js
assert.equal(titleCase("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast");
```

# --seed--

## --seed-contents--

```js
function titleCase(title) {
  return title;
}
```

# --solutions--

```js
function titleCase(title) {
  return title
    .split(" ")
    .map(word => word.charAt(0).toUpperCase() + word.slice(1).toLowerCase())
    .join(" ");
}
```
