---
id: 691b559495c5cb5a37b9b484
title: "チャレンジ 124: 在庫の更新"
challengeType: 28
dashedName: challenge-124
---

# --description--

あなたのストアの在庫を表す2次元配列と、受け取った出荷を表す別の2次元配列が与えられたとき、更新された在庫を返してください。

- 配列の各要素は `[quantity, "item"]` の形式で、`quantity` は整数、`"item"` は文字列です。
- 出荷の中で一致するアイテムがあれば、その数量を在庫に加えて更新してください。
- 受け取ったアイテムが現在の在庫に存在しない場合は、新しいエントリーとして在庫の最後に追加してください。
- 在庫は元の順序のまま返し、新しいアイテムは出荷に現れた順で最後に追加してください。

例えば、在庫が `[[2, "apples"], [5, "bananas"]]` で出荷が `[[1, "apples"], [3, "bananas"]]` の場合、 `[[3, "apples"], [8, "bananas"]]` を返します。

# --hints--

`updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]])` は `[[3, "apples"], [8, "bananas"]]` を返すはずです。

```js
assert.deepEqual(updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]]), [[3, "apples"], [8, "bananas"]]);
```

`updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]])` は `[[3, "apples"], [8, "bananas"], [4, "oranges"]]` を返すはずです。

```js
assert.deepEqual(updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]]), [[3, "apples"], [8, "bananas"], [4, "oranges"]]);
```

`updateInventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]])` は `[[10, "apples"], [30, "bananas"], [20, "oranges"]]` を返すはずです。

```js
assert.deepEqual(updateInventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]]), [[10, "apples"], [30, "bananas"], [20, "oranges"]]);
```

`updateInventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]])` は `[[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]]` を返すはずです。

```js
assert.deepEqual(updateInventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]]), [[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]]);
```

# --seed--

## --seed-contents--

```js
function updateInventory(inventory, shipment) {
  return inventory;
}
```

# --solutions--

```js
function updateInventory(inventory, shipment) {
  const inventoryMap = new Map();
  inventory.forEach(([qty, item], index) => {
    inventoryMap.set(item, index);
  });

  shipment.forEach(([qty, item]) => {
    if (inventoryMap.has(item)) {
      const index = inventoryMap.get(item);
      inventory[index][0] += qty;
    } else {
      inventory.push([qty, item]);
      inventoryMap.set(item, inventory.length - 1);
    }
  });

  return inventory;
}
```
