---
id: 691b559495c5cb5a37b9b480
title: "チャレンジ120：ポンドからキログラムへ"
challengeType: 28
dashedName: challenge-120
---

# --description--

ポンドで表された重さの数値が与えられたら、文字列`"(lbs) pounds equals (kgs) kilograms."`を返してください。

- `"(lbs)"`を入力された数値に置き換えてください。
- `"(kgs)"`をキログラムに換算した値に置き換えてください。小数点以下2桁に丸め、常に2桁の小数を含めてください。
- 1ポンドは0.453592キログラムです。
- 入力が`1`の場合は、`"pound"`の代わりに`"pounds"`を使ってください。
- 換算後の値が`1`の場合は、`"kilogram"`の代わりに`"kilograms"`を使ってください。

# --hints--

`convertToKgs(1)`は`"1 pound equals 0.45 kilograms."`を返すべきです。

```js
assert.equal(convertToKgs(1), "1 pound equals 0.45 kilograms.");
```

`convertToKgs(0)`は`"0 pounds equals 0.00 kilograms."`を返すべきです。

```js
assert.equal(convertToKgs(0), "0 pounds equals 0.00 kilograms.");
```

`convertToKgs(100)`は`"100 pounds equals 45.36 kilograms."`を返すべきです。

```js
assert.equal(convertToKgs(100), "100 pounds equals 45.36 kilograms.");
```

`convertToKgs(2.5)`は`"2.5 pounds equals 1.13 kilograms."`を返すべきです。

```js
assert.equal(convertToKgs(2.5), "2.5 pounds equals 1.13 kilograms.");
```

`convertToKgs(2.20462)`は`"2.20462 pounds equals 1.00 kilogram."`を返すべきです。

```js
assert.equal(convertToKgs(2.20462), "2.20462 pounds equals 1.00 kilogram.");
```

# --seed--

## --seed-contents--

```js
function convertToKgs(lbs) {
  return lbs;
}
```

# --solutions--

```js
function convertToKgs(lbs) {
  const KG_PER_POUND = 0.453592;
  const kgs = (lbs * KG_PER_POUND).toFixed(2);

  const poundWord = lbs === 1 ? "pound" : "pounds";
  const kilogramWord = kgs === "1.00" ? "kilogram" : "kilograms";

  return `${lbs} ${poundWord} equals ${kgs} ${kilogramWord}.`;
}
```
