---
id: 691b559495c5cb5a37b9b487
title: "チャレンジ 127: スピードチェック"
challengeType: 28
dashedName: challenge-127
---

# --description--

時速マイル（MPH）での走行速度と、時速キロメートル（KPH）での速度制限が与えられたとき、速度超過しているかどうか、警告を受けるか罰金を受けるかを判定してください。

- 1マイルは1.60934キロメートルです。
- 速度制限以下で走行している場合は `"Not Speeding"` を返してください。
- 速度制限を5KPH以下で超過している場合は `"Warning"` を返してください。
- 速度制限を5KPH以上超過している場合は `"Ticket"` を返してください。

# --hints--

`speedCheck(30, 70)` は `"Not Speeding"` を返すべきです。

```js
assert.equal(speedCheck(30, 70), "Not Speeding");
```

`speedCheck(40, 60)` は `"Warning"` を返すべきです。

```js
assert.equal(speedCheck(40, 60), "Warning");
```

`speedCheck(40, 65)` は `"Not Speeding"` を返すべきです。

```js
assert.equal(speedCheck(40, 65), "Not Speeding");
```

`speedCheck(60, 90)` は `"Ticket"` を返すべきです。

```js
assert.equal(speedCheck(60, 90), "Ticket");
```

`speedCheck(65, 100)` は `"Warning"` を返すべきです。

```js
assert.equal(speedCheck(65, 100), "Warning");
```

`speedCheck(88, 40)` は `"Ticket"` を返すべきです。

```js
assert.equal(speedCheck(88, 40), "Ticket");
```

# --seed--

## --seed-contents--

```js
function speedCheck(speedMph, speedLimitKph) {
  return speedMph;
}
```

# --solutions--

```js
function speedCheck(speedMph, speedLimitKph) {
  const speedKph = speedMph * 1.60934;

  if (speedKph <= speedLimitKph) {
    return "Not Speeding";
  }

  const over = speedKph - speedLimitKph;

  if (over <= 5) {
    return "Warning";
  }

  return "Ticket";
}
```
