---
id: 6925e2068081f40f549ced1c
title: "チャレンジ 138: 約数の合計"
challengeType: 28
dashedName: challenge-138
---

# --description--

正の整数が与えられたとき、その約数のすべての合計を返してください。

- 約数とは、その数を割り切る整数のことです（余りは`0`です）。
- 各約数は一度だけ数えます。

例えば、`6`が与えられた場合、`12`を返してください。なぜなら`6`の約数は`1`、`2`、`3`、`6`であり、それらの合計が`12`だからです。

# --hints--

`sumDivisors(6)`は`12`を返すべきです。

```js
assert.equal(sumDivisors(6), 12);
```

`sumDivisors(13)`は`14`を返すべきです。

```js
assert.equal(sumDivisors(13), 14);
```

`sumDivisors(28)`は`56`を返すべきです。

```js
assert.equal(sumDivisors(28), 56);
```

`sumDivisors(84)`は`224`を返すべきです。

```js
assert.equal(sumDivisors(84), 224);
```

`sumDivisors(549)`は`806`を返すべきです。

```js
assert.equal(sumDivisors(549), 806);
```

`sumDivisors(9348)`は`23520`を返すべきです。

```js
assert.equal(sumDivisors(9348), 23520);
```

# --seed--

## --seed-contents--

```js
function sumDivisors(n) {

  return n;
}
```

# --solutions--

```js
function sumDivisors(n) {
  let sum = 0;
  for (let i = 1; i <= n; i++) {
    if (n % i === 0) {
      sum += i;
    }
  }
  return sum;
}
```
