---
id: 6925e2068081f40f549ced1a
title: "チャレンジ 136: Markdown 画像パーサー"
challengeType: 28
dashedName: challenge-136
---

# --description--

Markdown の画像を表す文字列が与えられたとき、それに対応する HTML の文字列を返してください。

Markdown の画像は次の形式です：`"![alt text](image_url)"`。ここで：

- `alt text` は画像の説明（`alt` 属性の値）です。
- `image_url` は画像のソース URL（`src` 属性の値）です。

HTML の `img` タグの文字列を返し、`src` を画像の URL に、`alt` を代替テキストに設定してください。

例えば、`"![Cute cat](cat.png)"` が与えられたら `'<img src="cat.png" alt="Cute cat">'` を返してください。

- タグ、属性の順序、スペース、引用符の使い方が上記と同じであることを確認してください。

注意：コンソールは文字列内の HTML タグを表示しない場合があります。タグを含むログを確認するにはブラウザのコンソールをチェックしてください。

# --hints--

`parseImage("![Cute cat](cat.png)")` は `'<img src="cat.png" alt="Cute cat">'` を返すべきです。

```js
assert.equal(parseImage("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">');
```

`parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` は `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'` を返すべきです。

```js
assert.equal(parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">');
```

`parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` は `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'` を返すべきです。

```js
assert.equal(parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">');
```

# --seed--

## --seed-contents--

```js
function parseImage(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseImage(markdown) {
  const regex = /!\[(.*?)\]\((.*?)\)/;
  const match = markdown.match(regex);
  if (!match) return "";
  const alt = match[1];
  const src = match[2];
  return `<img src="${src}" alt="${alt}">`;
}
```
