---
id: 6925e2068081f40f549ced1d
title: "チャレンジ 139: じゃんけん"
challengeType: 28
dashedName: challenge-139
---

# --description--

2つの文字列が与えられます。最初の文字列はプレイヤー1を、2番目の文字列はプレイヤー2を表します。じゃんけんの勝者を判定してください。

- 入力される文字列は常に`"Rock"`、`"Paper"`、または`"Scissors"`のいずれかです。
- `"Rock"`は`"Scissors"`に勝ちます。
- `"Paper"`は`"Rock"`に勝ちます。
- `"Scissors"`は`"Paper"`に勝ちます。

返す値は以下の通りです：

- プレイヤー1が勝った場合は`"Player 1 wins"`を返してください。
- プレイヤー2が勝った場合は`"Player 2 wins"`を返してください。
- 両方のプレイヤーが同じ選択をした場合は`"Tie"`を返してください。

# --hints--

`rockPaperScissors("Rock", "Rock")`は`"Tie"`を返すべきです。

```js
assert.equal(rockPaperScissors("Rock", "Rock"), "Tie");
```

`rockPaperScissors("Rock", "Paper")`は`"Player 2 wins"`を返すべきです。

```js
assert.equal(rockPaperScissors("Rock", "Paper"), "Player 2 wins");
```

`rockPaperScissors("Scissors", "Paper")`は`"Player 1 wins"`を返すべきです。

```js
assert.equal(rockPaperScissors("Scissors", "Paper"), "Player 1 wins");
```

`rockPaperScissors("Rock", "Scissors")`は`"Player 1 wins"`を返すべきです。

```js
assert.equal(rockPaperScissors("Rock", "Scissors"), "Player 1 wins");
```

`rockPaperScissors("Scissors", "Scissors")`は`"Tie"`を返すべきです。

```js
assert.equal(rockPaperScissors("Scissors", "Scissors"), "Tie");
```

`rockPaperScissors("Scissors", "Rock")`は`"Player 2 wins"`を返すべきです。

```js
assert.equal(rockPaperScissors("Scissors", "Rock"), "Player 2 wins");
```

# --seed--

## --seed-contents--

```js
function rockPaperScissors(player1, player2) {

  return player1;
}
```

# --solutions--

```js
function rockPaperScissors(player1, player2) {
  if (player1 === player2) return "Tie";

  if (
    (player1 === "Rock" && player2 === "Scissors") ||
    (player1 === "Paper" && player2 === "Rock") ||
    (player1 === "Scissors" && player2 === "Paper")
  ) {
    return "Player 1 wins";
  } else {
    return "Player 2 wins";
  }
}
```
