---
id: 6925e2068081f40f549ced1b
title: "チャレンジ 137: Snowflake Generator"
challengeType: 28
dashedName: challenge-137
---

# --description--

改行文字 (`\n`) を使って改行を表現している複数行の文字列が与えられたとき、各行を水平方向に鏡像反転し、元の行の末尾に付け加えた新しい文字列を返してください。

- 行を鏡像反転するとは、スペースを含むすべての文字を逆順に並べることです。

例えば、`"* \n *\n* "` が与えられた場合、コンソールには次のように表示されます。

```sh
* 
 *
* 
```

返すべき文字列は `"*  *\n ** \n*  *"` で、コンソールには次のように表示されます。

```sh
*  *
 ** 
*  *
```

与えられた文字列と返す文字列の空白に注意してください。空白を削除しないようにしてください。

# --hints--

`generateSnowflake("* \n *\n* ")` は `"*  *\n ** \n*  *"` を返すべきです。

```js
assert.equal(generateSnowflake("* \n *\n* "), "*  *\n ** \n*  *");
```

`generateSnowflake("X=~")` は `"X=~~=X"` を返すべきです。

```js
assert.equal(generateSnowflake("X=~"), "X=~~=X");
```

`generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  ")` は `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "` を返すべきです。

```js
assert.equal(generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  "), " X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X ");
```

`generateSnowflake("*   *\n * * \n* * *\n * * \n*   *")` は `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"` を返すべきです。

```js
assert.equal(generateSnowflake("*   *\n * * \n* * *\n * * \n*   *"), "*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *");
```

`generateSnowflake("*  -\n * -\n*  -")` は `"*  --  *\n * -- * \n*  --  *"` を返すべきです。

```js
assert.equal(generateSnowflake("*  -\n * -\n*  -"), "*  --  *\n * -- * \n*  --  *");
```

# --seed--

## --seed-contents--

```js
function generateSnowflake(crystals) {

  return crystals;
}
```

# --solutions--

```js
function generateSnowflake(crystals) {
  const lines = crystals.split("\n");

  const mirroredLines = lines.map(line => {
    const reversed = line.split("").reverse().join("");
    return line + reversed;
  });

  return mirroredLines.join("\n");
}
```
