---
id: 69272dcf1c24b44fd79137c3
title: "チャレンジ 140: SCREAMING_SNAKE_CASE"
challengeType: 28
dashedName: challenge-140
---

# --description--

変数名を表す文字列が与えられたとき、その変数名を SCREAMING_SNAKE_CASE に変換して返してください。

与えられる変数名は次のいずれかの形式で書かれています：

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

これらの形式では、単語はアンダースコア（`_`）、ハイフン（`-`）、または新しい単語が大文字で始まることで区切られています。

SCREAMING_SNAKE_CASE に変換するには：

- すべての文字を大文字にする
- 単語をアンダースコア（`_`）で区切る

# --hints--

`toScreamingSnakeCase("userEmail")` は `"USER_EMAIL"` を返すべきです。

```js
assert.equal(toScreamingSnakeCase("userEmail"), "USER_EMAIL");
```

`toScreamingSnakeCase("UserPassword")` は `"USER_PASSWORD"` を返すべきです。

```js
assert.equal(toScreamingSnakeCase("UserPassword"), "USER_PASSWORD");
```

`toScreamingSnakeCase("user_id")` は `"USER_ID"` を返すべきです。

```js
assert.equal(toScreamingSnakeCase("user_id"), "USER_ID");
```

`toScreamingSnakeCase("user-address")` は `"USER_ADDRESS"` を返すべきです。

```js
assert.equal(toScreamingSnakeCase("user-address"), "USER_ADDRESS");
```

`toScreamingSnakeCase("username")` は `"USERNAME"` を返すべきです。

```js
assert.equal(toScreamingSnakeCase("username"), "USERNAME");
```

`toScreamingSnakeCase("my_variable_name")` は `"MY_VARIABLE_NAME"` を返すべきです。

```js
assert.equal(toScreamingSnakeCase("my_variable_name"), "MY_VARIABLE_NAME");
```

# --seed--

## --seed-contents--

```js
function toScreamingSnakeCase(variableName) {

  return variableName;
}
```

# --solutions--

```js
function toScreamingSnakeCase(variableName) {
  let temp = variableName.replace(/[-_]+/g, ' ');
  temp = temp.replace(/([a-z0-9])([A-Z])/g, '$1 $2');
  const words = temp.trim().split(/\s+/);
  return words.join('_').toUpperCase();
}
```
