---
id: 69272dcf1c24b44fd79137c4
title: "チャレンジ 141: 離陸用燃料"
challengeType: 28
dashedName: challenge-141
---

# --description--

現在の飛行機の燃料ガロン数と、目的地に到達するために必要な燃料リットル数が与えられたとき、追加で何ガロンの燃料を補給すべきかを求めてください。

- 1ガロンは3.78541リットルです。
- 飛行機にすでに十分な燃料がある場合は、`0`を返してください。
- 補給できるのは整数のガロンのみです。
- 返す数値に小数点は含めないでください。

# --hints--

`fuelToAdd(0, 1)`は`1`を返すべきです。

```js
assert.equal(fuelToAdd(0, 1), 1);
```

`fuelToAdd(5, 40)`は`6`を返すべきです。

```js
assert.equal(fuelToAdd(5, 40), 6);
```

`fuelToAdd(10, 30)`は`0`を返すべきです。

```js
assert.equal(fuelToAdd(10, 30), 0);
```

`fuelToAdd(896, 20500)`は`4520`を返すべきです。

```js
assert.equal(fuelToAdd(896, 20500), 4520);
```

`fuelToAdd(1000, 50000)`は`12209`を返すべきです。

```js
assert.equal(fuelToAdd(1000, 50000), 12209);
```

# --seed--

## --seed-contents--

```js
function fuelToAdd(currentGallons, requiredLiters) {

  return currentGallons;
}
```

# --solutions--

```js
function fuelToAdd(currentGallons, requiredLiters) {
  const litersPerGallon = 3.78541;
  const currentLiters = currentGallons * litersPerGallon;
  if (currentLiters >= requiredLiters) return 0;
  const litersNeeded = requiredLiters - currentLiters;
  return Math.ceil(litersNeeded / litersPerGallon);
}
```
