---
id: 69272dcf1c24b44fd79137c5
title: "チャレンジ142：文字列の合計を求める"
challengeType: 28
dashedName: challenge-142
---

# --description--

数字とその他の文字が混在する文字列が与えられたとき、その文字列内のすべての数字の合計を返してください。

- 連続する数字は1つの数値として扱います。たとえば、`"13"` は 1 + 3 ではなく 13 とカウントします。
- 数字以外の文字は無視してください。

# --hints--

`stringSum("3apples2bananas")` は `5` を返すべきです。

```js
assert.equal(stringSum("3apples2bananas"), 5);
```

`stringSum("10cats5dogs2birds")` は `17` を返すべきです。

```js
assert.equal(stringSum("10cats5dogs2birds"), 17);
```

`stringSum("125344")` は `125344` を返すべきです。

```js
assert.equal(stringSum("125344"), 125344);
```

`stringSum("a1b20c300")` は `321` を返すべきです。

```js
assert.equal(stringSum("a1b20c300"), 321);
```

`stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` は `1653` を返すべきです。

```js
assert.equal(stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653);
```

# --seed--

## --seed-contents--

```js
function stringSum(str) {

  return str;
}
```

# --solutions--

```js
function stringSum(str) {
  const matches = str.match(/\d+/g); // find all consecutive digits
  return matches.reduce((sum, num) => sum + Number(num), 0);
}
```
