---
id: 69306364df283fcaff2e1ad5
title: "チャレンジ 144: 解像度の連続達成"
challengeType: 28
dashedName: challenge-144
---

# --description--

配列の配列が与えられます。各サブ配列はその日の解像度の活動を表し、3つの整数が含まれています：その日に歩いたステップ数、その日のスクリーンタイムの分数、その日に読んだページ数です。あなたが解像度を守っているかどうかを判定してください。

- 最初のサブ配列は1日目、2番目は2日目、というように続きます。

以下の3つの目標すべてを満たした日を成功した日とみなします：

- 少なくとも10,000ステップ歩いた。
- スクリーンタイムは120分以下だった。
- 少なくとも5ページ読んだ。

すべての日が成功している場合、`"Resolution on track: N day streak."` を返してください。ここで`N`は成功した日の数です。

1日以上成功していない日がある場合、`"Resolution failed on day X: N day streak."` を返してください。ここで`X`は最初に成功しなかった日の番号、`N`は最初に成功しなかった日より前の成功した日の数です。

# --hints--

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]])` は `"Resolution on track: 3 day streak."` を返すべきです

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]]), "Resolution on track: 3 day streak.");
```

`resolutionStreak([[10000, 120, 5], [10950, 121, 11]])` は `"Resolution failed on day 2: 1 day streak."` を返すべきです

```js
assert.equal(resolutionStreak([[10000, 120, 5], [10950, 121, 11]]), "Resolution failed on day 2: 1 day streak.");
```

`resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]])` は `"Resolution failed on day 3: 2 day streak."` を返すべきです

```js
assert.equal(resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]]), "Resolution failed on day 3: 2 day streak.");
```

`resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]])` は `"Resolution failed on day 4: 3 day streak."` を返すべきです

```js
assert.equal(resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]]), "Resolution failed on day 4: 3 day streak.");
```

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]])` は `"Resolution on track: 10 day streak."` を返すべきです

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]]), "Resolution on track: 10 day streak.");
```

# --seed--

## --seed-contents--

```js
function resolutionStreak(days) {

  return days;
}
```

# --solutions--

```js
function resolutionStreak(days) {
  let streak = 0;

  for (let i = 0; i < days.length; i++) {
    const [steps, screen, pages] = days[i];
    if (steps >= 10000 && screen <= 120 && pages >= 5) {
      streak++;
    } else {
      return `Resolution failed on day ${i + 1}: ${streak} day streak.`;
    }
  }

  return `Resolution on track: ${streak} day streak.`;
}
```
