---
id: 69306364df283fcaff2e1ad8
title: "チャレンジ 147: うるう年計算機"
challengeType: 28
dashedName: challenge-147
---

# --description--

整数の年が与えられたとき、それがうるう年かどうかを判定してください。

うるう年とは、以下のルールを満たす年です：

- その年が4で割り切れること、かつ
- その年が100で割り切れないこと、ただし
- その年が400で割り切れる場合は例外としてうるう年とする。

# --hints--

`isLeapYear(2024)` は `true` を返すべきです。

```js
assert.isTrue(isLeapYear(2024));
```

`isLeapYear(2023)` は `false` を返すべきです。

```js
assert.isFalse(isLeapYear(2023));
```

`isLeapYear(2100)` は `false` を返すべきです。

```js
assert.isFalse(isLeapYear(2100));
```

`isLeapYear(2000)` は `true` を返すべきです。

```js
assert.isTrue(isLeapYear(2000));
```

`isLeapYear(1999)` は `false` を返すべきです。

```js
assert.isFalse(isLeapYear(1999));
```

`isLeapYear(2040)` は `true` を返すべきです。

```js
assert.isTrue(isLeapYear(2040));
```

`isLeapYear(2026)` は `false` を返すべきです。

```js
assert.isFalse(isLeapYear(2026));
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

  return year;
}
```

# --solutions--

```js
function isLeapYear(year) {
    return ((year % 4 === 0 && year % 100 !== 0) || (year % 400 === 0))
}
```
