---
id: 69306364df283fcaff2e1ad7
title: "チャレンジ146：テーブルの左利き用席"
challengeType: 28
dashedName: challenge-146
---

# --description--

4x2の行列（配列の配列）で表された食事の席配置が与えられたとき、左利きの人が座れる席の数を判定してください。

- 左利きの人は、自分のすぐ左隣の席に右利きの人がいる場所には座れません。

与えられた行列では：

- `"R"` は右利きの人が座っている席です。
- `"L"` は左利きの人が座っている席です。
- `"U"` は空席です。
- 座れるのは空席のみです。
- 上の行の席は「下向き」、下の行の席は「上向き」（テーブルのように）なので、左右は席の向きに対して相対的です。
- 角の席は隣に1席しかありません。

例えば、与えられた行列では：

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

左上の席は、左隣に右利きの人がいるため座れません。その他の2つの空席は、左隣に右利きの人がいないので座れます。

# --hints--

`findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` は `2` を返すべきです。

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2);
```

`findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` は `8` を返すべきです。

```js
assert.equal(findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8);
```

`findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` は `0` を返すべきです。

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0);
```

`findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` は `1` を返すべきです。

```js
assert.equal(findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1);
```

`findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` は `5` を返すべきです。

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5);
```

# --seed--

## --seed-contents--

```js
function findLeftHandedSeats(table) {

  return table;
}
```

# --solutions--

```js
function findLeftHandedSeats(table) {
  let availableSeats = 0;

  const [topRow, bottomRow] = table;

  for (let i=0; i<topRow.length; i++) {
    if (topRow[i] === 'U') {
      if (i === 3) {
        availableSeats++;
      } else if (topRow[i+1] !== 'R') {
        availableSeats++;
      }
    }
  }

  for (let i=0; i<bottomRow.length; i++) {
    if (bottomRow[i] === 'U') {
      if (i === 0) {
        availableSeats++;
      } else if (bottomRow[i-1] !== 'R') {
        availableSeats++;
      }
    }
  }

  return availableSeats;
}
```
