---
id: 69306364df283fcaff2e1ada
title: "チャレンジ149: vOwElcAsE"
challengeType: 28
dashedName: challenge-149
---

# --description--

文字列が与えられたとき、すべての母音を大文字に変換し、それ以外のアルファベット文字を小文字に変換した新しい文字列を返してください。

- 母音はどのケースでも `"a"`、`"e"`、`"i"`、`"o"`、`"u"` です。
- アルファベット以外の文字は変更しないでください。

# --hints--

`vowelCase("vowelcase")` は `"vOwElcAsE"` を返すべきです。

```js
assert.equal(vowelCase("vowelcase"), "vOwElcAsE");
```

`vowelCase("coding is fun")` は `"cOdIng Is fUn"` を返すべきです。

```js
assert.equal(vowelCase("coding is fun"), "cOdIng Is fUn");
```

`vowelCase("HELLO, world!")` は `"hEllO, wOrld!"` を返すべきです。

```js
assert.equal(vowelCase("HELLO, world!"), "hEllO, wOrld!");
```

`vowelCase("git cherry-pick")` は `"gIt chErry-pIck"` を返すべきです。

```js
assert.equal(vowelCase("git cherry-pick"), "gIt chErry-pIck");
```

`vowelCase("HEAD~1")` は `"hEAd~1"` を返すべきです。

```js
assert.equal(vowelCase("HEAD~1"), "hEAd~1");
```

# --seed--

## --seed-contents--

```js
function vowelCase(str) {

  return str;
}
```

# --solutions--

```js
function vowelCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (vowels.includes(char)) {
      result += char.toUpperCase();
    } else if (/[a-zA-Z]/.test(char)) {
      result += char.toLowerCase();
    } else {
      result += char;
    }
  }

  return result;
}
```
