---
id: 69306364df283fcaff2e1ad6
title: "チャレンジ 145: N番目のフィボナッチ数"
challengeType: 28
dashedName: challenge-145
---

# --description--

整数 `n` が与えられたとき、フィボナッチシーケンスの `n` 番目の数を返してください。

フィボナッチシーケンスは、各数が直前の2つの数の合計になる数列です。シーケンスの最初の10個の数は `0`、`1`、`1`、`2`、`3`、`5`、`8`、`13`、`21`、`34` です。

# --hints--

`nthFibonacci(4)` は `2` を返すべきです。

```js
assert.equal(nthFibonacci(4), 2);
```

`nthFibonacci(10)` は `34` を返すべきです。

```js
assert.equal(nthFibonacci(10), 34);
```

`nthFibonacci(15)` は `377` を返すべきです。

```js
assert.equal(nthFibonacci(15), 377);
```

`nthFibonacci(40)` は `63245986` を返すべきです。

```js
assert.equal(nthFibonacci(40), 63245986);
```

`nthFibonacci(75)` は `1304969544928657` を返すべきです。

```js
assert.equal(nthFibonacci(75), 1304969544928657);
```

# --seed--

## --seed-contents--

```js
function nthFibonacci(n) {

  return n;
}
```

# --solutions--

```js
function nthFibonacci(n) {
  if (n === 1) return 0;
  if (n === 2) return 1;

  let a = 0;
  let b = 1;

  for (let i = 3; i <= n; i++) {
    const next = a + b;
    a = b;
    b = next;
  }

  return b;
}
```
