---
id: 69306364df283fcaff2e1ad9
title: "チャレンジ 148: タイヤの空気圧"
challengeType: 28
dashedName: challenge-148
---

# --description--

車両の4つのタイヤの空気圧をpsi単位で表す4つの数字の配列と、タイヤの最小および最大空気圧をbar単位で表す2つの数字の配列が与えられたとき、各タイヤの状態を表す4つの文字列の配列を返してください。

- 1 bar は14.5038 psiに相当します。

各タイヤについて、次の値を含む配列を返してください。

- 空気圧が許容される最小値を下回っている場合は `"Low"`。
- 空気圧が最小値と最大値の間にある場合は `"Good"`。
- 空気圧が最大値を超えている場合は `"High"`。

# --hints--

`tireStatus([32, 28, 35, 29], [2, 3])` は `["Good", "Low", "Good", "Low"]` を返すべきです。

```js
assert.deepEqual(tireStatus([32, 28, 35, 29], [2, 3]), ["Good", "Low", "Good", "Low"]);
```

`tireStatus([32, 28, 35, 30], [2, 2.3])` は `["Good", "Low", "High", "Good"]` を返すべきです。

```js
assert.deepEqual(tireStatus([32, 28, 35, 30], [2, 2.3]), ["Good", "Low", "High", "Good"]);
```

`tireStatus([29, 26, 31, 28], [2.1, 2.5])` は `["Low", "Low", "Good", "Low"]` を返すべきです。

```js
assert.deepEqual(tireStatus([29, 26, 31, 28], [2.1, 2.5]), ["Low", "Low", "Good", "Low"]);
```

`tireStatus([31, 31, 30, 29], [1.5, 2])` は `["High", "High", "High", "Good"]` を返すべきです。

```js
assert.deepEqual(tireStatus([31, 31, 30, 29], [1.5, 2]), ["High", "High", "High", "Good"]);
```

`tireStatus([30, 28, 30, 29], [1.9, 2.1])` は `["Good", "Good", "Good", "Good"]` を返すべきです。

```js
assert.deepEqual(tireStatus([30, 28, 30, 29], [1.9, 2.1]), ["Good", "Good", "Good", "Good"]);
```

# --seed--

## --seed-contents--

```js
function tireStatus(pressuresPSI, rangeBar) {

  return pressuresPSI;
}
```

# --solutions--

```js
function tireStatus(pressuresPSI, rangeBar) {
  const psiToBar = 1 / 14.5038;
  const [minBar, maxBar] = rangeBar;

  return pressuresPSI.map(psi => {
    const bar = psi * psiToBar;
    if (bar < minBar) return "Low";
    if (bar > maxBar) return "High";
    return "Good";
  });
}
```
