---
id: 69373793f5a867f769cde139
title: "チャレンジ 154: ホールのパー"
challengeType: 28
dashedName: challenge-154
---

# --description--

2つの整数、ゴルフのホールのパーとそのホールでゴルファーが打ったストローク数が与えられたとき、ゴルフ用語を使ってゴルファーのスコアを返してください。

返す値:

- ストロークが1回なら `"Hole in one!"` を返してください。
- パーより2打少なければ `"Eagle"` を返してください。
- パーより1打少なければ `"Birdie"` を返してください。
- パーと同じストローク数なら `"Par"` を返してください。
- パーより1打多ければ `"Bogey"` を返してください。
- パーより2打多ければ `"Double bogey"` を返してください。

# --hints--

`golfScore(3, 3)` は `"Par"` を返すべきです。

```js
assert.equal(golfScore(3, 3), "Par");
```

`golfScore(4, 3)` は `"Birdie"` を返すべきです。

```js
assert.equal(golfScore(4, 3), "Birdie");
```

`golfScore(3, 1)` は `"Hole in one!"` を返すべきです。

```js
assert.equal(golfScore(3, 1), "Hole in one!");
```

`golfScore(5, 7)` は `"Double bogey"` を返すべきです。

```js
assert.equal(golfScore(5, 7), "Double bogey");
```

`golfScore(4, 5)` は `"Bogey"` を返すべきです。

```js
assert.equal(golfScore(4, 5), "Bogey");
```

`golfScore(5, 3)` は `"Eagle"` を返すべきです。

```js
assert.equal(golfScore(5, 3), "Eagle");
```

# --seed--

## --seed-contents--

```js
function golfScore(par, strokes) {

  return par;
}
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) return "Hole in one!";

  const diff = strokes - par;

  if (diff === -2) return "Eagle";
  if (diff === -1) return "Birdie";
  if (diff === 0) return "Par";
  if (diff === 1) return "Bogey";
  if (diff === 2) return "Double bogey";
}
```
