---
id: 69373793f5a867f769cde138
title: "チャレンジ 153: 三目並べ"
challengeType: 28
dashedName: challenge-153
---

# --description--

3×3の行列（配列の配列）で表された完成した三目並べのゲームが与えられたとき、勝者を判定してください。

- 与えられた行列の各要素は`"X"`か`"O"`のいずれかです。

プレイヤーは、横、縦、または斜めに自分の文字が3つ並んでいれば勝ちです。

返す値：

- プレイヤーXが3つ並んでいれば`"X wins"`を返してください。
- プレイヤーOが3つ並んでいれば`"O wins"`を返してください。
- どちらのプレイヤーも3つ並んでいなければ`"Draw"`を返してください。

# --hints--

`ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]])`は`"X wins"`を返すべきです。

```js
assert.equal(ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]]), "X wins");
```

`ticTacToe([["O", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]])`は`"X wins"`を返すべきです。

```js
assert.equal(ticTacToe([["O", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]])`は`"Draw"`を返すべきです。

```js
assert.equal(ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]]), "Draw");
```

`ticTacToe([["X", "X", "O"], ["X", "O", "O"], ["O", "O", "X"]])`は`"O wins"`を返すべきです。

```js
assert.equal(ticTacToe([["X", "X", "O"], ["X", "O", "O"], ["O", "O", "X"]]), "O wins");
```

`ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]])`は`"X wins"`を返すべきです。

```js
assert.equal(ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]])`は`"Draw"`を返すべきです。

```js
assert.equal(ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]]), "Draw");
```

# --seed--

## --seed-contents--

```js
function ticTacToe(board) {

  return board;
}
```

# --solutions--

```js
function ticTacToe(board) {
  const lines = [];

  for (let row of board) {
    lines.push(row);
  }

  for (let c = 0; c < 3; c++) {
    lines.push([board[0][c], board[1][c], board[2][c]]);
  }

  lines.push([board[0][0], board[1][1], board[2][2]]);
  lines.push([board[0][2], board[1][1], board[2][0]]);

  for (let line of lines) {
    if (line.every(cell => cell === "X")) return "X wins";
    if (line.every(cell => cell === "O")) return "O wins";
  }

  return "Draw";
}
```
