---
id: 69373793f5a867f769cde136
title: "チャレンジ 151: ソートされた配列？"
challengeType: 28
dashedName: challenge-151
---

# --description--

数値の配列が与えられたとき、その数値が昇順、降順、またはどちらでもないかを判定してください。

与えられた配列が以下の場合：

- 昇順（最小から最大）なら、`"Ascending"`を返してください。
- 降順（最大から最小）なら、`"Descending"`を返してください。
- 昇順でも降順でもない場合は、`"Not sorted"`を返してください。

# --hints--

`isSorted([1, 2, 3, 4, 5])`は`"Ascending"`を返すべきです。

```js
assert.equal(isSorted([1, 2, 3, 4, 5]), "Ascending");
```

`isSorted([10, 8, 6, 4, 2])`は`"Descending"`を返すべきです。

```js
assert.equal(isSorted([10, 8, 6, 4, 2]), "Descending");
```

`isSorted([1, 3, 2, 4, 5])`は`"Not sorted"`を返すべきです。

```js
assert.equal(isSorted([1, 3, 2, 4, 5]), "Not sorted");
```

`isSorted([3.14, 2.71, 1.61, 0.57])`は`"Descending"`を返すべきです。

```js
assert.equal(isSorted([3.14, 2.71, 1.61, 0.57]), "Descending");
```

`isSorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9])`は`"Ascending"`を返すべきです。

```js
assert.equal(isSorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9]), "Ascending");
```

`isSorted([0.4, 0.5, 0.3])`は`"Not sorted"`を返すべきです。

```js
assert.equal(isSorted([0.4, 0.5, 0.3]), "Not sorted");
```

# --seed--

## --seed-contents--

```js
function isSorted(arr) {

  return arr;
}
```

# --solutions--

```js
function isSorted(arr) {
  let ascending = true;
  let descending = true;

  for (let i = 1; i < arr.length; i++) {
    if (arr[i] < arr[i - 1]) ascending = false;
    if (arr[i] > arr[i - 1]) descending = false;
  }

  if (ascending) return "Ascending";
  if (descending) return "Descending";
  return "Not sorted";
}
```
