---
id: 69373793f5a867f769cde135
title: "チャレンジ150：Markdownの番号なしリストパーサー"
challengeType: 28
dashedName: challenge-150
---

# --description--

有効なMarkdownの番号なしリストの文字列が与えられたとき、それに対応するHTMLの文字列を返してください。

番号なしリストは1つ以上のリスト項目で構成されます。有効なリスト項目は独立した行に現れ、以下の条件を満たします。

- ダッシュ（`"-"`）で始まり、
- その後に少なくとも1つのスペースがあり、
- その後にリスト項目のテキストが続きます。

リストは改行文字（`"\n"`）で区切られた単一の文字列として与えられます。項目のテキストに改行文字は含めないでください。

各リスト項目をHTMLの`li`タグで囲み、リスト全体を`ul`タグで囲んでください。

例えば、`"- Item A\n- Item B"`が与えられたら、`"<ul><li>Item A</li><li>Item B</li></ul>"`を返します。

注意：コンソールは文字列内のHTMLタグをログに表示しない場合があります。タグを含むログを確認するにはブラウザのコンソールをチェックしてください。

# --hints--

`parseUnorderedList("- Item A\n- Item B")`は`"<ul><li>Item A</li><li>Item B</li></ul>"`を返すはずです。

```js
assert.equal(parseUnorderedList("- Item A\n- Item B"), "<ul><li>Item A</li><li>Item B</li></ul>");
```

`parseUnorderedList("-  JavaScript\n-  Python")`は`"<ul><li>JavaScript</li><li>Python</li></ul>"`を返すはずです。

```js
assert.equal(parseUnorderedList("-  JavaScript\n-  Python"), "<ul><li>JavaScript</li><li>Python</li></ul>");
```

`parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla")`は`"<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>"`を返すはずです。

```js
assert.equal(parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla"), "<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>");
```

`parseUnorderedList("- A-1\n- A-2\n- B-1")`は`"<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>"`を返すはずです。

```js
assert.equal(parseUnorderedList("- A-1\n- A-2\n- B-1"), "<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>");
```

# --seed--

## --seed-contents--

```js
function parseUnorderedList(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseUnorderedList(markdown) {
  const lines = markdown.split("\n");
  const listItems = [];
  
  for (let line of lines) {
    listItems.push(`<li>${line.slice(1).trim()}</li>`);
  }
  
  return `<ul>${listItems.join("")}</ul>`;
}
```
