---
id: 69373793f5a867f769cde137
title: "チャレンジ 152: 循環素数"
challengeType: 28
dashedName: challenge-152
---

# --description--

整数が与えられたとき、それが循環素数かどうかを判定してください。

循環素数とは、その数字のすべての回転が素数である整数のことです。

例えば、`197` は循環素数です。なぜなら、その数字のすべての回転である `197`、`971`、および `719` がすべて素数だからです。

# --hints--

`isCircularPrime(197)` は `true` を返すべきです。

```js
assert.isTrue(isCircularPrime(197));
```

`isCircularPrime(23)` は `false` を返すべきです。

```js
assert.isFalse(isCircularPrime(23));
```

`isCircularPrime(13)` は `true` を返すべきです。

```js
assert.isTrue(isCircularPrime(13));
```

`isCircularPrime(89)` は `false` を返すべきです。

```js
assert.isFalse(isCircularPrime(89));
```

`isCircularPrime(1193)` は `true` を返すべきです。

```js
assert.isTrue(isCircularPrime(1193));
```

# --seed--

## --seed-contents--

```js
function isCircularPrime(n) {

  return n;
}
```

# --solutions--

```js
function isPrime(n) {
  if (n < 2) return false;
  for (let i = 2, sqrt = Math.floor(Math.sqrt(n)); i <= sqrt; i++) {
    if (n % i === 0) return false;
  }
  return true;
}

function rotations(num) {
  const str = num.toString();
  const result = [];
  for (let i = 0; i < str.length; i++) {
    result.push(str.slice(i) + str.slice(0, i));
  }
  return result.map(Number);
}

function isCircularPrime(num) {
  const rots = rotations(num);
  return rots.every(isPrime);
}
```
