---
id: 69373793f5a867f769cde13a
title: "チャレンジ 155: 作物を植える"
challengeType: 28
dashedName: challenge-155
---

# --description--

農場の畑のサイズを表す整数と、農場の畑のサイズの単位を表す`"acres"`または`"hectares"`、および作物の種類が与えられたとき、その畑にその種類の植物をいくつ植えられるかを求めてください。

- 1エーカーは4046.86平方メートルです。
- 1ヘクタールは10,000平方メートルです。

以下は入力として与えられる作物のリストと、1つの植物が必要とするスペースです：

|作物|1株あたりのスペース|
|-|-|
|`"corn"`|1平方メートル|
|`"wheat"`|0.1平方メートル|
|`"soybeans"`|0.5平方メートル|
|`"tomatoes"`|0.25平方メートル|
|`"lettuce"`|0.2平方メートル|

畑に植えられる植物の数を、小数点以下切り捨てで返してください。

# --hints--

`getNumberOfPlants(1, "acres", "corn")`は`4046`を返すべきです。

```js
assert.equal(getNumberOfPlants(1, "acres", "corn"), 4046);
```

`getNumberOfPlants(2, "hectares", "lettuce")`は`100000`を返すべきです。

```js
assert.equal(getNumberOfPlants(2, "hectares", "lettuce"), 100000);
```

`getNumberOfPlants(20, "acres", "soybeans")`は`161874`を返すべきです。

```js
assert.equal(getNumberOfPlants(20, "acres", "soybeans"), 161874);
```

`getNumberOfPlants(3.75, "hectares", "tomatoes")`は`150000`を返すべきです。

```js
assert.equal(getNumberOfPlants(3.75, "hectares", "tomatoes"), 150000);
```

`getNumberOfPlants(16.75, "acres", "tomatoes")`は`271139`を返すべきです。

```js
assert.equal(getNumberOfPlants(16.75, "acres", "tomatoes"), 271139);
```

# --seed--

## --seed-contents--

```js
function getNumberOfPlants(fieldSize, unit, crop) {

  return fieldSize;
}
```

# --solutions--

```js
function getNumberOfPlants(fieldSize, unit, crop) {
  const cropSpace = {
    corn: 1,
    wheat: 0.1,
    soybeans: 0.5,
    tomatoes: 0.25,
    lettuce: 0.2
  };

  let fieldInSqM = unit === "acres" ? fieldSize * 4046.86 : fieldSize * 10000;

  return Math.floor(fieldInSqM / cropSpace[crop.toLowerCase()]);
}
```
