---
id: 6939b873185d8e00d453563d
title: "チャレンジ 158: 配列の入れ替え"
challengeType: 28
dashedName: challenge-158
---

# --description--

2つの値を持つ配列が与えられたとき、値を入れ替えた配列を返してください。

例えば、`["A", "B"]` が与えられたら `["B", "A"]` を返します。

# --hints--

`arraySwap(["A", "B"])` は `["B", "A"]` を返すはずです。

```js
assert.deepEqual(arraySwap(["A", "B"]), ["B", "A"]);
```

`arraySwap([25, 20])` は `[20, 25]` を返すはずです。

```js
assert.deepEqual(arraySwap([25, 20]), [20, 25]);
```

`arraySwap([true, false])` は `[false, true]` を返すはずです。

```js
assert.deepEqual(arraySwap([true, false]), [false, true]);
```

`arraySwap(["1", 1])` は `[1, "1"]` を返すはずです。

```js
assert.deepEqual(arraySwap(["1", 1]), [1, "1"]);
```

# --seed--

## --seed-contents--

```js
function arraySwap(arr) {

  return arr;
}
```

# --solutions--

```js
function arraySwap(arr) {

  return [arr[1], arr[0]];
}
```
