---
id: 6939b873185d8e00d453563b
title: "チャレンジ 156: 奇数か偶数か？"
challengeType: 28
dashedName: challenge-156
---

# --description--

正の整数が与えられたら、奇数の場合は `"Odd"` を、偶数の場合は `"Even"` を返してください。

# --hints--

`oddOrEven(1)` は `"Odd"` を返すべきです。

```js
assert.equal(oddOrEven(1), "Odd");
```

`oddOrEven(2)` は `"Even"` を返すべきです。

```js
assert.equal(oddOrEven(2), "Even");
```

`oddOrEven(13)` は `"Odd"` を返すべきです。

```js
assert.equal(oddOrEven(13), "Odd");
```

`oddOrEven(196)` は `"Even"` を返すべきです。

```js
assert.equal(oddOrEven(196), "Even");
```

`oddOrEven(123456789)` は `"Odd"` を返すべきです。

```js
assert.equal(oddOrEven(123456789), "Odd");
```

# --seed--

## --seed-contents--

```js
function oddOrEven(n) {

  return n;
}
```

# --solutions--

```js
function oddOrEven(n) {
  return n % 2 === 0 ? "Even" : "Odd";
}
```
