---
id: 6939b873185d8e00d453563e
title: "チャレンジ159：整数の斜辺"
challengeType: 28
dashedName: challenge-159
---

# --description--

2つの正の整数が直角三角形の2つの脚（短い2辺）の長さを表すとき、斜辺の長さが整数かどうかを判定してください。

斜辺の長さは、2つの脚の長さの2乗を足し合わせ、その合計の平方根を取ることで計算されます（a<sup>2</sup> + b<sup>2</sup> = c<sup>2</sup>）。

# --hints--

`isIntegerHypotenuse(3, 4)` は `true` を返す必要があります。

```js
assert.isTrue(isIntegerHypotenuse(3, 4));
```

`isIntegerHypotenuse(2, 3)` は `false` を返す必要があります。

```js
assert.isFalse(isIntegerHypotenuse(2, 3));
```

`isIntegerHypotenuse(5, 12)` は `true` を返す必要があります。

```js
assert.isTrue(isIntegerHypotenuse(5, 12));
```

`isIntegerHypotenuse(10, 10)` は `false` を返す必要があります。

```js
assert.isFalse(isIntegerHypotenuse(10, 10));
```

`isIntegerHypotenuse(780, 1040)` は `true` を返す必要があります。

```js
assert.isTrue(isIntegerHypotenuse(780, 1040));
```

`isIntegerHypotenuse(250, 333)` は `false` を返す必要があります。

```js
assert.isFalse(isIntegerHypotenuse(250, 333));
```

# --seed--

## --seed-contents--

```js
function isIntegerHypotenuse(a, b) {

  return a;
}
```

# --solutions--

```js
function isIntegerHypotenuse(a, b) {
  const sum = a * a + b * b;
  const c = Math.floor(Math.sqrt(sum));
  return c * c === sum;
}
```
