---
id: 6939b873185d8e00d453563f
title: "チャレンジ 160: ナイトの動き"
challengeType: 28
dashedName: challenge-160
---

# --description--

チェス盤上のナイトの位置が与えられたとき、ナイトが移動できる有効なマスの数を返してください。

標準的なチェス盤は8×8で、列は左から右に`A`から`H`まで、行は下から上に`1`から`8`までラベル付けされています。以下のようになっています：

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

ナイトは「L字型」に動きます：一方向（横または縦）に2マス、そして直角方向に1マス動きます。

つまり、ナイトは最大で8つの位置に移動できますが、盤の端に近い場合は移動可能なマスが少なくなります。例えば、ナイトが`A1`にいる場合、移動できるのは`B3`か`C2`だけです。

# --hints--

`knightMoves("A1")`は`2`を返すべきです。

```js
assert.equal(knightMoves("A1"), 2);
```

`knightMoves("D4")`は`8`を返すべきです。

```js
assert.equal(knightMoves("D4"), 8);
```

`knightMoves("G6")`は`6`を返すべきです。

```js
assert.equal(knightMoves("G6"), 6);
```

`knightMoves("B8")`は`3`を返すべきです。

```js
assert.equal(knightMoves("B8"), 3);
```

`knightMoves("H3")`は`4`を返すべきです。

```js
assert.equal(knightMoves("H3"), 4);
```

# --seed--

## --seed-contents--

```js
function knightMoves(position) {

  return position;
}
```

# --solutions--

```js
function knightMoves(position) {
  const col = position[0].charCodeAt(0) - 65;
  const row = parseInt(position[1]) - 1;

  const moves = [
    [2, 1], [2, -1],
    [-2, 1], [-2, -1],
    [1, 2], [1, -2],
    [-1, 2], [-1, -2]
  ];

  let valid = 0;

  for (const [dx, dy] of moves) {
    const newCol = col + dx;
    const newRow = row + dy;

    if (newCol >= 0 && newCol < 8 && newRow >= 0 && newRow < 8) {
      valid++;
    }
  }

  return valid;
}
```
