---
id: 6939b873185d8e00d453563c
title: "チャレンジ157：Markdownリンクパーサー"
challengeType: 28
dashedName: challenge-157
---

# --description--

Markdownのリンクの文字列が与えられたとき、それに対応するHTMLの文字列を返してください。

Markdownの画像は次の形式です：`"[link_text](link_url)"`。`a`タグの文字列を返し、`href`を`link_url`に設定し、`link_text`をタグの内容として使います。

例えば、`"[freeCodeCamp](https://freecodecamp.org/)"`が与えられたら`'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`を返します。

注意：コンソールは文字列内のHTMLタグをログに表示しない場合があります。タグを含むログを見るにはブラウザのコンソールを確認してください。

# --hints--

`parseLink("[freeCodeCamp](https://freecodecamp.org/)")`は`'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`を返すべきです。

```js
assert.equal(parseLink("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>');
```

`parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)")`は`'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'`を返すべきです。

```js
assert.equal(parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>');
```

`parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")`は`'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>'`を返すべきです。

```js
assert.equal(parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>');
```

# --seed--

## --seed-contents--

```js
function parseLink(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseLink(markdown) {
  const match = markdown.match(/\[(.*?)\]\((.*?)\)/);

  const [, text, url] = match;
  return `<a href="${url}">${text}</a>`;
}
```
