---
id: 694596b0585c11170ac7c7fa
title: "チャレンジ 162: エネルギー消費量"
challengeType: 28
dashedName: challenge-162
---

# --description--

運動中に消費したカロリー数と、その運動中に電子機器が使用したワット時数が与えられたとき、どちらがより多くのエネルギーを使ったかを判定してください。

比較するために、以下の換算を使って両方の値をジュールに変換してください。

- 1カロリーは4184ジュールに相当します。
- 1ワット時は3600ジュールに相当します。

返す値は以下の通りです。

- 運動の方が多くエネルギーを使った場合は `"Workout"` を返してください。
- デバイスの方が多くエネルギーを使った場合は `"Devices"` を返してください。
- 両方が同じ量のエネルギーを使った場合は `"Equal"` を返してください。

# --hints--

`compareEnergy(250, 50)` は `"Workout"` を返すべきです。

```js
assert.equal(compareEnergy(250, 50), "Workout");
```

`compareEnergy(100, 200)` は `"Devices"` を返すべきです。

```js
assert.equal(compareEnergy(100, 200), "Devices");
```

`compareEnergy(450, 523)` は `"Equal"` を返すべきです。

```js
assert.equal(compareEnergy(450, 523), "Equal");
```

`compareEnergy(300, 75)` は `"Workout"` を返すべきです。

```js
assert.equal(compareEnergy(300, 75), "Workout");
```

`compareEnergy(200, 250)` は `"Devices"` を返すべきです。

```js
assert.equal(compareEnergy(200, 250), "Devices");
```

`compareEnergy(900, 1046)` は `"Equal"` を返すべきです。

```js
assert.equal(compareEnergy(900, 1046), "Equal");
```

# --seed--

## --seed-contents--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {

  return caloriesBurned;
}
```

# --solutions--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {
  const workoutEnergy = caloriesBurned * 4184;
  const deviceEnergy = wattHoursUsed * 3600;

  console.log(workoutEnergy);
  console.log(deviceEnergy);
  if (workoutEnergy > deviceEnergy) return "Workout";
  if (deviceEnergy > workoutEnergy) return "Devices";
  return "Equal";
}
```
