---
id: 694596b0585c11170ac7c7fb
title: "チャレンジ 163: 子音ケース"
challengeType: 28
dashedName: challenge-163
---

# --description--

変数名を表す文字列が与えられたとき、次のルールに従って子音ケースに変換してください。

- すべての子音は大文字に変換します。
- すべての母音（`a`、`e`、`i`、`o`、`u`のいずれのケースも）は小文字に変換します。
- すべてのハイフン（`-`）はアンダースコア（`_`）に変換します。

# --hints--

`toConsonantCase("helloworld")`は`"HeLLoWoRLD"`を返すべきです。

```js
assert.equal(toConsonantCase("helloworld"), "HeLLoWoRLD");
```

`toConsonantCase("HELLOWORLD")`は`"HeLLoWoRLD"`を返すべきです。

```js
assert.equal(toConsonantCase("HELLOWORLD"), "HeLLoWoRLD");
```

`toConsonantCase("_hElLO-WOrlD-")`は`"_HeLLo_WoRLD_"`を返すべきです。

```js
assert.equal(toConsonantCase("_hElLO-WOrlD-"), "_HeLLo_WoRLD_");
```

`toConsonantCase("_~-generic_~-variable_~-name_~-here-~_")`は`"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"`を返すべきです。

```js
assert.equal(toConsonantCase("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_");
```

# --seed--

## --seed-contents--

```js
function toConsonantCase(str) {

  return str;
}
```

# --solutions--

```js
function toConsonantCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (char === "-") {
      result += "_";
    } else if (/[a-zA-Z]/.test(char)) {
      if (vowels.includes(char)) {
        result += char.toLowerCase();
      } else {
        result += char.toUpperCase();
      }
    } else {
      result += char;
    }
  }

  return result;
}
```
