---
id: 694596b0585c11170ac7c7fd
title: "チャレンジ 165: クラスの平均点"
challengeType: 28
dashedName: challenge-165
---

# --description--

試験の点数（数値）の配列が与えられたとき、以下のチャートに従って平均点を文字の成績で返してください。

| 平均点 | 成績 |
| - | - |
| 97-100 | `"A+"` |
| 93-96 | `"A"` |
| 90-92 | `"A-"` |
| 87-89 | `"B+"` |
| 83-86 | `"B"` |
| 80-82 | `"B-"` |
| 77-79 | `"C+"` |
| 73–76 | `"C"` |
| 70-72 | `"C-"` |
| 67-69 | `"D+"` |
| 63-66 | `"D"` |
| 60–62 | `"D-"` |
| 60未満 | `"F"` |

配列内のすべての点数を足して、点数の総数で割ることで平均点を計算してください。

# --hints--

`getAverageGrade([92, 91, 90, 94, 89, 93])` は `"A-"` を返すべきです。

```js
assert.equal(getAverageGrade([92, 91, 90, 94, 89, 93]), "A-");
```

`getAverageGrade([84, 89, 85, 100, 91, 88, 79])` は `"B+"` を返すべきです。

```js
assert.equal(getAverageGrade([84, 89, 85, 100, 91, 88, 79]), "B+");
```

`getAverageGrade([63, 69, 65, 66, 71, 64, 65])` は `"D"` を返すべきです。

```js
assert.equal(getAverageGrade([63, 69, 65, 66, 71, 64, 65]), "D");
```

`getAverageGrade([97, 98, 99, 100, 96, 97, 98, 99, 100])` は `"A+"` を返すべきです。

```js
assert.equal(getAverageGrade([97, 98, 99, 100, 96, 97, 98, 99, 100]), "A+");
```

`getAverageGrade([75, 100, 88, 79, 80, 78, 64, 60])` は `"C+"` を返すべきです。

```js
assert.equal(getAverageGrade([75, 100, 88, 79, 80, 78, 64, 60]), "C+");
```

`getAverageGrade([45, 48, 50, 52, 100, 54, 56, 58, 59])` は `"F"` を返すべきです。

```js
assert.equal(getAverageGrade([45, 48, 50, 52, 100, 54, 56, 58, 59]), "F");
```

# --seed--

## --seed-contents--

```js
function getAverageGrade(scores) {

  return scores;
}
```

# --solutions--

```js
function getAverageGrade(scores) {
  const avg =
    scores.reduce((sum, score) => sum + score, 0) / scores.length;

  if (avg >= 97) return "A+";
  if (avg >= 93) return "A";
  if (avg >= 90) return "A-";

  if (avg >= 87) return "B+";
  if (avg >= 83) return "B";
  if (avg >= 80) return "B-";

  if (avg >= 77) return "C+";
  if (avg >= 73) return "C";
  if (avg >= 70) return "C-";

  if (avg >= 67) return "D+";
  if (avg >= 63) return "D";
  if (avg >= 60) return "D-";

  return "F";
}
```
