---
id: 696655d24b614176d4c9b78c
title: "チャレンジ169：FizzBuzzミニ"
challengeType: 28
dashedName: challenge-169
---

# --description--

整数が与えられたら、次のルールに基づいて文字列を返してください。

- その数が3で割り切れる場合は `"Fizz"` を返してください。
- その数が5で割り切れる場合は `"Buzz"` を返してください。
- その数が3と5の両方で割り切れる場合は `"FizzBuzz"` を返してください。
- それ以外の場合は、与えられた数を文字列として返してください。

# --hints--

`fizzBuzzMini(3)` は `"Fizz"` を返すべきです。

```js
assert.equal(fizzBuzzMini(3), "Fizz");
```

`fizzBuzzMini(4)` は `"4"` を返すべきです。

```js
assert.strictEqual(fizzBuzzMini(4), "4");
```

`fizzBuzzMini(35)` は `"Buzz"` を返すべきです。

```js
assert.equal(fizzBuzzMini(35), "Buzz");
```

`fizzBuzzMini(75)` は `"FizzBuzz"` を返すべきです。

```js
assert.equal(fizzBuzzMini(75), "FizzBuzz");
```

`fizzBuzzMini(98)` は `"98"` を返すべきです。

```js
assert.strictEqual(fizzBuzzMini(98), "98");
```

# --seed--

## --seed-contents--

```js
function fizzBuzzMini(n) {

  return n;
}
```

# --solutions--

```js
function fizzBuzzMini(n) {
  if (n % 3 === 0 && n % 5 === 0) return "FizzBuzz";
  if (n % 3 === 0) return "Fizz";
  if (n % 5 === 0) return "Buzz";
  return String(n);
}
```
