---
id: 696655d24b614176d4c9b78b
title: "チャレンジ168：スケールされた画像"
challengeType: 28
dashedName: challenge-168
---

# --description--

画像の幅と高さを表す文字列と、画像をスケールする数値が与えられたとき、スケールされた幅と高さを返してください。

- 入力の文字列は`"WxH"`の形式です。例えば、`"800x600"`です。
- スケールは幅と高さに乗算する数値です。

スケールされた寸法を同じ`"WxH"`形式で返してください。

# --hints--

`scaleImage("800x600", 2)`は`"1600x1200"`を返すべきです。

```js
assert.equal(scaleImage("800x600", 2), "1600x1200");
```

`scaleImage("100x100", 10)`は`"1000x1000"`を返すべきです。

```js
assert.equal(scaleImage("100x100", 10), "1000x1000");
```

`scaleImage("1024x768", 0.5)`は`"512x384"`を返すべきです。

```js
assert.equal(scaleImage("1024x768", 0.5), "512x384");
```

`scaleImage("300x200", 1.5)`は`"450x300"`を返すべきです。

```js
assert.equal(scaleImage("300x200", 1.5), "450x300");
```

# --seed--

## --seed-contents--

```js
function scaleImage(size, scale) {

  return size;
}
```

# --solutions--

```js
function scaleImage(size, scale) {
  const [width, height] = size.split("x").map(Number);

  const newWidth = width * scale;
  const newHeight = height * scale;

  return `${newWidth}x${newHeight}`;
}
```
