---
id: 696655d24b614176d4c9b78a
title: "チャレンジ 167: ビンゴ！文字"
challengeType: 28
dashedName: challenge-167
---

# --description--

数字が与えられたら、それに対応するビンゴの文字（大文字）を返してください。ビンゴの数字は以下のようにグループ分けされています：

| 文字 | 数字の範囲 |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

# --hints--

`getBingoLetter(75)` は `"O"` を返す必要があります。

```js
assert.equal(getBingoLetter(75), "O");
```

`getBingoLetter(54)` は `"G"` を返す必要があります。

```js
assert.equal(getBingoLetter(54), "G");
```

`getBingoLetter(25)` は `"I"` を返す必要があります。

```js
assert.equal(getBingoLetter(25), "I");
```

`getBingoLetter(38)` は `"N"` を返す必要があります。

```js
assert.equal(getBingoLetter(38), "N");
```

`getBingoLetter(11)` は `"B"` を返す必要があります。

```js
assert.equal(getBingoLetter(11), "B");
```

# --seed--

## --seed-contents--

```js
function getBingoLetter(n) {

  return n;
}
```

# --solutions--

```js
function getBingoLetter(n) {
  if (n >= 1 && n <= 15) return "B";
  if (n >= 16 && n <= 30) return "I";
  if (n >= 31 && n <= 45) return "N";
  if (n >= 46 && n <= 60) return "G";
  if (n >= 61 && n <= 75) return "O";
  return null;
}
```
