---
id: 696655d24b614176d4c9b78d
title: "チャレンジ 170: 奇数日か偶数日か"
challengeType: 28
dashedName: challenge-170
---

# --description--

Unixエポックからのミリ秒数で表されたタイムスタンプが与えられたとき、次を返してください：

- そのタイムスタンプの日付が奇数日なら `"odd"` を返す。
- そのタイムスタンプの日付が偶数日なら `"even"` を返す。

例えば、`1769472000000`（2026年1月27日のタイムスタンプ）が与えられた場合、日付の数字（`"odd"`）が奇数なので `27` を返します。

注意：タイムスタンプはミリ秒単位で、ローカル時間ではなくUTCタイムゾーンの日付を使う必要があります。

# --hints--

`oddOrEvenDay(1769472000000)` は `"odd"` を返すべきです。

```js
assert.equal(oddOrEvenDay(1769472000000), "odd");
```

`oddOrEvenDay(1769444440000)` は `"even"` を返すべきです。

```js
assert.equal(oddOrEvenDay(1769444440000), "even");
```

`oddOrEvenDay(6739456780000)` は `"odd"` を返すべきです。

```js
assert.equal(oddOrEvenDay(6739456780000), "odd");
```

`oddOrEvenDay(1)` は `"odd"` を返すべきです。

```js
assert.equal(oddOrEvenDay(1), "odd");
```

`oddOrEvenDay(86400000)` は `"even"` を返すべきです。

```js
assert.equal(oddOrEvenDay(86400000), "even");
```

# --seed--

## --seed-contents--

```js
function oddOrEvenDay(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function oddOrEvenDay(timestamp) {
  const date = new Date(timestamp);
  const day = date.getUTCDate();

  return day % 2 === 0 ? "even" : "odd";
}
```
