---
id: 696655d24b614176d4c9b789
title: "チャレンジ 166: 16進数バリデーター"
challengeType: 28
dashedName: challenge-166
---

# --description--

文字列が有効なCSSの16進数カラーかどうかを判定してください。有効なCSSの16進数カラーは以下の条件を満たす必要があります：

- `#`で始まり、
- その後に3文字または6文字の16進数文字が続くこと。

16進数文字は数字の`0`から`9`および文字の`a`から`f`（大文字・小文字は区別しません）です。

# --hints--

`isValidHex("#123")`は`true`を返すべきです。

```js
assert.isTrue(isValidHex("#123"));
```

`isValidHex("#123abc")`は`true`を返すべきです。

```js
assert.isTrue(isValidHex("#123abc"));
```

`isValidHex("#ABCDEF")`は`true`を返すべきです。

```js
assert.isTrue(isValidHex("#ABCDEF"));
```

`isValidHex("#0a1B2c")`は`true`を返すべきです。

```js
assert.isTrue(isValidHex("#0a1B2c"));
```

`isValidHex("#12G")`は`false`を返すべきです。

```js
assert.isFalse(isValidHex("#12G"));
```

`isValidHex("#1234567")`は`false`を返すべきです。

```js
assert.isFalse(isValidHex("#1234567"));
```

`isValidHex("#12 3")`は`false`を返すべきです。

```js
assert.isFalse(isValidHex("#12 3"));
```

`isValidHex("fff")`は`false`を返すべきです。

```js
assert.isFalse(isValidHex("fff"));
```

# --seed--

## --seed-contents--

```js
function isValidHex(str) {

  return str;
}
```

# --solutions--

```js
function isValidHex(str) {
  if (str[0] !== "#") return false;

  const hex = str.slice(1);
  if (hex.length !== 3 && hex.length !== 6) return false;

  for (let char of hex) {
    if (!/[0-9a-fA-F]/.test(char)) {
      return false;
    }
  }

  return true;
}
```
