---
id: 69738771fb5a7b8b24cca2a3
title: "チャレンジ 177: 文字列ミラー"
challengeType: 28
dashedName: challenge-177
---

# --description--

文字列が与えられたとき、その文字列の末尾に逆順にしたコピーを付け加えた新しい文字列を返してください。

# --hints--

`mirror("freeCodeCamp")` は `"freeCodeCamppmaCedoCeerf"` を返すべきです。

```js
assert.equal(mirror("freeCodeCamp"), "freeCodeCamppmaCedoCeerf");
```

`mirror("RaceCar")` は `"RaceCarraCecaR"` を返すべきです。

```js
assert.equal(mirror("RaceCar"), "RaceCarraCecaR");
```

`mirror("helloworld")` は `"helloworlddlrowolleh"` を返すべきです。

```js
assert.equal(mirror("helloworld"), "helloworlddlrowolleh");
```

`mirror("The quick brown fox...")` は `"The quick brown fox......xof nworb kciuq ehT"` を返すべきです。

```js
assert.equal(mirror("The quick brown fox..."), "The quick brown fox......xof nworb kciuq ehT");
```

# --seed--

## --seed-contents--

```js
function mirror(str) {

  return str;
}
```

# --solutions--

```js
function mirror(str) {
  const reversed = str.split("").reverse().join("");
  return str + reversed;
}
```
