---
id: 69738771fb5a7b8b24cca2a4
title: "チャレンジ 178: 文字列を切り詰める"
challengeType: 28
dashedName: challenge-178
---

# --description--

文字列が20文字以下の場合は、そのまま返してください。20文字を超える場合は、最初の17文字に切り詰め、末尾に`"..."`を追加して（合計20文字にして）結果を返してください。

# --hints--

`truncateText("Hello, world!")` は `"Hello, world!"` を返すべきです。

```js
assert.equal(truncateText("Hello, world!"), "Hello, world!");
```

`truncateText("This string should get truncated.")` は `"This string shoul..."` を返すべきです。

```js
assert.equal(truncateText("This string should get truncated."), "This string shoul...");
```

`truncateText("Exactly twenty chars")` は `"Exactly twenty chars"` を返すべきです。

```js
assert.equal(truncateText("Exactly twenty chars"), "Exactly twenty chars");
```

`truncateText(".....................")` は `"...................."` を返すべきです。

```js
assert.equal(truncateText("....................."), "....................");
```

# --seed--

## --seed-contents--

```js
function truncateText(text) {

  return text;
}
```

# --solutions--

```js
function truncateText(text) {
  if (text.length <= 20) return text;
  return text.slice(0, 17) + "...";
}
```
