---
id: 69738771fb5a7b8b24cca29e
title: "チャレンジ172: 文字と数字"
challengeType: 28
dashedName: challenge-172
---

# --description--

文字と数字だけを含む文字列が与えられたとき、文字から数字、または数字から文字に切り替わるたびにハイフン（`-`）を挿入した新しい文字列を返してください。

# --hints--

`separateLettersAndNumbers("ABC123")` は `"ABC-123"` を返すべきです。

```js
assert.equal(separateLettersAndNumbers("ABC123"), "ABC-123");
```

`separateLettersAndNumbers("Route66")` は `"Route-66"` を返すべきです。

```js
assert.equal(separateLettersAndNumbers("Route66"), "Route-66");
```

`separateLettersAndNumbers("H3LL0W0RLD")` は `"H-3-LL-0-W-0-RLD"` を返すべきです。

```js
assert.equal(separateLettersAndNumbers("H3LL0W0RLD"), "H-3-LL-0-W-0-RLD");
```

`separateLettersAndNumbers("a1b2c3d4")` は `"a-1-b-2-c-3-d-4"` を返すべきです。

```js
assert.equal(separateLettersAndNumbers("a1b2c3d4"), "a-1-b-2-c-3-d-4");
```

# --seed--

## --seed-contents--

```js
function separateLettersAndNumbers(str) {

  return str;
}
```

# --solutions--

```js
function separateLettersAndNumbers(str) {
  if (str.length === 0) return str;

  let result = str[0];

  for (let i = 1; i < str.length; i++) {
    const prev = str[i - 1];
    const curr = str[i];

    const prevIsLetter = /[a-zA-Z]/.test(prev);
    const currIsLetter = /[a-zA-Z]/.test(curr);

    if (prevIsLetter !== currIsLetter) {
      result += "-";
    }

    result += curr;
  }

  return result;
}
```
