---
id: 69738771fb5a7b8b24cca2a1
title: "チャレンジ175: デジタルデトックス"
challengeType: 28
dashedName: challenge-175
---

# --description--

ログイン履歴の配列が与えられたとき、デジタルデトックスの目標を達成しているかどうかを判定してください。

各ログは `"YYYY-MM-DD HH:mm:ss"` の形式の文字列です。

次の両方の条件を満たしている場合、デジタルデトックスの目標を達成しています。

- どの4時間の期間内でもログインは1回以下であること。
- どの日でもログインは2回以下であること。

# --hints--

`digitalDetox(["2026-02-01 08:00:00", "2026-02-01 12:30:00"])` は `true` を返すべきです。

```js
assert.isTrue(digitalDetox(["2026-02-01 08:00:00", "2026-02-01 12:30:00"]));
```

`digitalDetox(["2026-02-01 04:00:00", "2026-02-01 07:30:00"])` は `false` を返すべきです。

```js
assert.isFalse(digitalDetox(["2026-02-01 04:00:00", "2026-02-01 07:30:00"]));
```

`digitalDetox(["2026-01-31 08:21:30", "2026-01-31 14:30:00", "2026-02-01 08:00:00", "2026-02-01 12:30:00"])` は `true` を返すべきです。

```js
assert.isTrue(digitalDetox(["2026-01-31 08:21:30", "2026-01-31 14:30:00", "2026-02-01 08:00:00", "2026-02-01 12:30:00"]));
```

`digitalDetox(["2026-01-31 10:40:21", "2026-01-31 15:19:41", "2026-01-31 21:49:50", "2026-02-01 09:30:00"])` は `false` を返すべきです。

```js
assert.isFalse(digitalDetox(["2026-01-31 10:40:21", "2026-01-31 15:19:41", "2026-01-31 21:49:50", "2026-02-01 09:30:00"]));
```

`digitalDetox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 22:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 09:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"])` は `true` を返すべきです。

```js
assert.isTrue(digitalDetox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 20:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 09:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"]));
```

`digitalDetox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 22:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 01:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"])` は `false` を返すべきです。

```js
assert.isFalse(digitalDetox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 22:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 01:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"]));
```

# --seed--

## --seed-contents--

```js
function digitalDetox(logs) {

  return logs;
}
```

# --solutions--

```js
function digitalDetox(logs) {
  if (logs.length === 0) return true;

  const times = logs
    .map(log => new Date(log.replace(" ", "T") + "Z").getTime())
    .sort((a, b) => a - b);

  const ONE_DAY = 24 * 60 * 60 * 1000;
  const FOUR_HOURS = 4 * 60 * 60 * 1000;

  for (let i = 1; i < times.length; i++) {
    if (times[i] - times[i - 1] < FOUR_HOURS) {
      return false;
    }
  }

  const dailyCounts = {};
  for (const time of times) {
    const day = new Date(time).toISOString().slice(0, 10);
    dailyCounts[day] = (dailyCounts[day] || 0) + 1;
    if (dailyCounts[day] > 2) {
      return false;
    }
  }

  return true;
}
```
