---
id: 69738771fb5a7b8b24cca2a0
title: "チャレンジ 174: 星座ファインダー"
challengeType: 28
dashedName: challenge-174
---

# --description--

`"YYYY-MM-DD"` 形式の日付文字列が与えられたとき、以下のチャートを使ってその日の星座を返してください。

| 日付範囲 | 星座 |
| - | - |
| 3月21日 - 4月19日 | `"Aries"` |
| 4月20日 - 5月20日 | `"Taurus"` |
| 5月21日 - 6月20日 | `"Gemini"` |
| 6月21日 - 7月22日 | `"Cancer"` |
| 7月23日 - 8月22日 | `"Leo"` |
| 8月23日 - 9月22日 | `"Virgo"` |
| 9月23日 - 10月22日 | `"Libra"` |
| 10月23日 - 11月21日 | `"Scorpio"` |
| 11月22日 - 12月21日 | `"Sagittarius"` |
| 12月22日 - 1月19日 | `"Capricorn"` |
| 1月20日 - 2月18日 | `"Aquarius"` |
| 2月19日 - 3月20日 | `"Pisces"` |

- 星座は月と日だけで決まるため、年は無視してかまいません。

# --hints--

`getSign("2026-01-31")` は `"Aquarius"` を返すべきです。

```js
assert.equal(getSign("2026-01-31"), "Aquarius");
```

`getSign("2001-06-10")` は `"Gemini"` を返すべきです。

```js
assert.equal(getSign("2001-06-10"), "Gemini");
```

`getSign("1985-09-07")` は `"Virgo"` を返すべきです。

```js
assert.equal(getSign("1985-09-07"), "Virgo");
```

`getSign("2023-03-19")` は `"Pisces"` を返すべきです。

```js
assert.equal(getSign("2023-03-19"), "Pisces");
```

`getSign("2045-11-05")` は `"Scorpio"` を返すべきです。

```js
assert.equal(getSign("2045-11-05"), "Scorpio");
```

`getSign("1985-12-06")` は `"Sagittarius"` を返すべきです。

```js
assert.equal(getSign("1985-12-06"), "Sagittarius");
```

`getSign("2025-12-30")` は `"Capricorn"` を返すべきです。

```js
assert.equal(getSign("2025-12-30"), "Capricorn");
```

`getSign("2018-10-08")` は `"Libra"` を返すべきです。

```js
assert.equal(getSign("2018-10-08"), "Libra");
```

`getSign("1958-05-04")` は `"Taurus"` を返すべきです。

```js
assert.equal(getSign("1958-05-04"), "Taurus");
```

# --seed--

## --seed-contents--

```js
function getSign(dateStr) {

  return dateStr;
}
```

# --solutions--

```js
function getSign(dateStr) {
  const [, monthStr, dayStr] = dateStr.split("-");
  const month = Number(monthStr);
  const day = Number(dayStr);

  if ((month === 3 && day >= 21) || (month === 4 && day <= 19)) return "Aries";
  if ((month === 4 && day >= 20) || (month === 5 && day <= 20)) return "Taurus";
  if ((month === 5 && day >= 21) || (month === 6 && day <= 20)) return "Gemini";
  if ((month === 6 && day >= 21) || (month === 7 && day <= 22)) return "Cancer";
  if ((month === 7 && day >= 23) || (month === 8 && day <= 22)) return "Leo";
  if ((month === 8 && day >= 23) || (month === 9 && day <= 22)) return "Virgo";
  if ((month === 9 && day >= 23) || (month === 10 && day <= 22)) return "Libra";
  if ((month === 10 && day >= 23) || (month === 11 && day <= 21)) return "Scorpio";
  if ((month === 11 && day >= 22) || (month === 12 && day <= 21)) return "Sagittarius";
  if ((month === 12 && day >= 22) || (month === 1 && day <= 19)) return "Capricorn";
  if ((month === 1 && day >= 20) || (month === 2 && day <= 18)) return "Aquarius";
  if ((month === 2 && day >= 19) || (month === 3 && day <= 20)) return "Pisces";
}
```
