---
id: 69738771fb5a7b8b24cca29f
title: "チャレンジ 173: 有効なポーンの動き"
challengeType: 28
dashedName: challenge-173
---

# --description--

チェスボード上の自分のポーンの位置が与えられたとき、そのポーンが移動できるすべての有効なマスを昇順で配列として返してください。

標準的なチェスボードは8×8で、列は左から右へ`A`から`H`まで、行は下から上へ`1`から`8`までラベル付けされています。見た目は次のようになります：

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

このチャレンジでは：

- あなたはボードの下側のプレイヤーです。
- ポーンは「上」に1マスだけ動けます。
- ただし、ポーンが初期の行（2行目）にいる場合は、1マスまたは2マス上に動けます。

例えば、`"D4"`が与えられたら、ポーンが動ける唯一のマスである`["D5"]`を返してください。
`"B2"`が与えられたら、初期行にいるため、昇順で`["B3", "B4"]`を返してください。

# --hints--

`findPawnMoves("D4")`は`["D5"]`を返すべきです。

```js
assert.deepEqual(findPawnMoves("D4"), ["D5"]);
```

`findPawnMoves("B2")`は`["B3", "B4"]`を返すべきです。

```js
assert.deepEqual(findPawnMoves("B2"), ["B3", "B4"]);
```

`findPawnMoves("A7")`は`["A8"]`を返すべきです。

```js
assert.deepEqual(findPawnMoves("A7"), ["A8"]);
```

`findPawnMoves("G2")`は`["G3", "G4"]`を返すべきです。

```js
assert.deepEqual(findPawnMoves("G2"), ["G3", "G4"]);
```

`findPawnMoves("E3")`は`["E4"]`を返すべきです。

```js
assert.deepEqual(findPawnMoves("E3"), ["E4"]);
```

# --seed--

## --seed-contents--

```js
function findPawnMoves(position) {

  return position;
}

```

# --solutions--

```js
function findPawnMoves(position) {
  const column = position[0].toUpperCase();
  const row = parseInt(position[1]);
  const moves = [];

  if (row >= 8) return moves;

  moves.push(`${column}${row + 1}`);

  if (row === 2) moves.push(`${column}${row + 2}`);

  return moves;
}
```
