---
id: 69738771fb5a7b8b24cca2a5
title: "チャレンジ179: ポケットの小銭"
challengeType: 28
dashedName: challenge-179
---

# --description--

ポケットの中のコインを表す整数の配列が与えられます。各整数はコインの価値をセント単位で表しています。合計金額を `"$D.CC"` の形式で返してください。

- 100セントは1ドルに相当します。
- 戻り値では、1ドル未満の金額には先頭にゼロを付け、セントは常に2桁で表示してください。

# --hints--

`countChange([25, 10, 5, 1])` は `"$0.41"` を返すべきです。

```js
assert.equal(countChange([25, 10, 5, 1]), "$0.41");
```

`countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25])` は `"$1.43"` を返すべきです。

```js
assert.equal(countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25]), "$1.43");
```

`countChange([100, 25, 100, 1000, 5, 500, 2000, 25])` は `"$37.55"` を返すべきです。

```js
assert.equal(countChange([100, 25, 100, 1000, 5, 500, 2000, 25]), "$37.55");
```

`countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10])` は `"$0.70"` を返すべきです。

```js
assert.equal(countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10]), "$0.70");
```

`countChange([1])` は `"$0.01"` を返すべきです。

```js
assert.equal(countChange([1]), "$0.01");
```

`countChange([25, 25, 25, 25])` は `"$1.00"` を返すべきです。

```js
assert.equal(countChange([25, 25, 25, 25]), "$1.00");
```

# --seed--

## --seed-contents--

```js
function countChange(change) {

  return change;
}
```

# --solutions--

```js
function countChange(change) {
  const totalCents = change.reduce((sum, coin) => sum + coin, 0);

  const dollars = Math.floor(totalCents / 100);
  const cents = totalCents % 100;

  return `$${dollars}.${cents.toString().padStart(2, "0")}`;
}
```
