---
id: 697a49e9860d24853adef67d
title: "チャレンジ 192: 2026年冬季大会 13日目: ノルディック複合"
challengeType: 28
dashedName: challenge-192
---

# --description--

選手のジャンプスコアの配列が与えられたとき、ノルディック複合のクロスカントリー部分のスタート遅延時間を計算してください。

ジャンプスコアが最も高い選手が最初にスタートします（遅延時間は0秒）。他の選手は、ジャンプスコアが最も高い選手との差に応じて遅れてスタートします。

各選手の遅延時間を計算するには、最も高いジャンプスコアからその選手のジャンプスコアを引き、その結果に1.5を乗算してください。遅延時間は最も近い整数に切り上げます。

# --hints--

`calculateStartDelays([120, 110, 125])` は `[8, 23, 0]` を返すべきです。

```js
assert.deepEqual(calculateStartDelays([120, 110, 125]), [8, 23, 0]);
```

`calculateStartDelays([118, 125, 122, 120])` は `[11, 0, 5, 8]` を返すべきです。

```js
assert.deepEqual(calculateStartDelays([118, 125, 122, 120]), [11, 0, 5, 8]);
```

`calculateStartDelays([100, 105, 95, 110, 120, 115, 108])` は `[30, 23, 38, 15, 0, 8, 18]` を返すべきです。

```js
assert.deepEqual(calculateStartDelays([100, 105, 95, 110, 120, 115, 108]), [30, 23, 38, 15, 0, 8, 18]);
```

`calculateStartDelays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124])` は `[3, 11, 6, 18, 21, 15, 8, 26, 0, 12]` を返すべきです。

```js
assert.deepEqual(calculateStartDelays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124]), [3, 11, 6, 18, 21, 15, 8, 26, 0, 12]);
```

# --seed--

## --seed-contents--

```js
function calculateStartDelays(jumpScores) {

  return jumpScores
}
```

# --solutions--

```js
function calculateStartDelays(jumpScores) {
  const bestJump = Math.max(...jumpScores);
  return jumpScores.map(score => Math.ceil((bestJump - score) * 1.5));
}
```
