---
id: 697a49e6ff50d756c9b69365
title: "チャレンジ 188: 2026年冬季大会 第9日目: スケルトン"
challengeType: 28
dashedName: challenge-188
---

# --description--

スケルトントラックのカーブを表す文字列が与えられたとき、そのトラックの難易度を判定してください。

- 与えられる文字列は以下の文字のみで構成されます：
  - `"L"` は左カーブを表します
  - `"R"` は右カーブを表します
  - `"S"` は直線区間を表します

- 方向転換がある場合は15点加算されます（`"L"`の後に`"R"`が続く、またはその逆の場合）。
- その他のカーブはそれぞれ5点加算されます（その他の`"L"`または`"R"`の文字）。
- 直線区間は0点です。

トラックの難易度は合計得点に基づきます。以下を返してください：

- 合計が0～100の場合は`"Easy"`
- 合計が101～200の場合は`"Medium"`
- 合計が200を超える場合は`"Hard"`

# --hints--

`getDifficulty("SLSLLSRRLSRLRL")` は`"Easy"`を返すべきです。

```js
assert.equal(getDifficulty("SLSLLSRRLSRLRL"), "Easy");
```

`getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` は`"Hard"`を返すべきです。

```js
assert.equal(getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard");
```

`getDifficulty("SRRRRLSLLRLRSSRLSRL")` は`"Medium"`を返すべきです。

```js
assert.equal(getDifficulty("SRRRRLSLLRLRSSRLSRL"), "Medium");
```

`getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` は`"Hard"`を返すべきです。

```js
assert.equal(getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard");
```

`getDifficulty("SLLSSLRLSLSLRSLSSLRL")` は`"Medium"`を返すべきです。

```js
assert.equal(getDifficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium");
```

`getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` は`"Easy"`を返すべきです。

```js
assert.equal(getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy");
```

# --seed--

## --seed-contents--

```js
function getDifficulty(track) {

  return track;
}
```

# --solutions--

```js
function getDifficulty(track) {
  let score = 0;

  for (let i = 0; i < track.length; i++) {
    const current = track[i];
    const previous = i > 0 ? track[i - 1] : null;

    if (current === 'S') {
      score += 0;
    } else if (current === 'L' || current === 'R') {
      if (previous && previous !== 'S' && previous !== current) {
        score += 15;
      } else if (previous !== 'S') {
        score += 5;
      } else {
        score += 5;
      }
    }
  }

  if (score <= 100) return 'Easy';
  if (score <= 200) return 'Medium';
  return 'Hard';
}
```
