---
id: 697a49e6ff50d756c9b69360
title: "チャレンジ 183: 2026年冬季大会 第4日目: スキージャンプ"
challengeType: 28
dashedName: challenge-183
---

# --description--

距離点、スタイル点、風補正値、K点ボーナス値が与えられたとき、スキージャンプの得点を計算し、メダルを獲得したかどうかを判定してください。

- 得点はこれら4つの値を合計して計算します。

他のジャンパーの現在の合計得点は以下の通りです：

```sh
165.5
172.0
158.0
180.0
169.5
175.0
162.0
170.0
```

- あなたの得点が最高なら、`"Gold"`を返してください。
- 2番目に高ければ、`"Silver"`を返してください。
- 3番目に高ければ、`"Bronze"`を返してください。
- それ以外の場合は、`"No Medal"`を返してください。

# --hints--

`skiJumpMedal(125.0, 58.0, 0.0, 6.0)`は`"Gold"`を返すべきです。

```js
assert.equal(skiJumpMedal(125.0, 58.0, 0.0, 6.0), "Gold");
```

`skiJumpMedal(119.0, 50.0, 1.0, 4.0)`は`"Bronze"`を返すべきです。

```js
assert.equal(skiJumpMedal(119.0, 50.0, 1.0, 4.0), "Bronze");
```

`skiJumpMedal(122.0, 52.0, -1.0, 4.0)`は`"Silver"`を返すべきです。

```js
assert.equal(skiJumpMedal(122.0, 52.0, -1.0, 4.0), "Silver");
```

`skiJumpMedal(118.0, 50.5, -1.5, 4.0)`は`"No Medal"`を返すべきです。

```js
assert.equal(skiJumpMedal(118.0, 50.5, -1.5, 4.0), "No Medal");
```

`skiJumpMedal(124.0, 50.5, 2.0, 5.0)`は`"Gold"`を返すべきです。

```js
assert.equal(skiJumpMedal(124.0, 50.5, 2.0, 5.0), "Gold");
```

`skiJumpMedal(119.0, 49.5, 0.0, 3.0)`は`"No Medal"`を返すべきです。

```js
assert.equal(skiJumpMedal(119.0, 49.5, 0.0, 3.0), "No Medal");
```

# --seed--

## --seed-contents--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {

  return distancePoints;
}
```

# --solutions--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {
  const myScore = distancePoints + stylePoints + windComp + kPointBonus;
  const otherScores = [165.5, 172.0, 158.0, 180.0, 169.5, 175.0, 162.0, 170.0];
  const allScores = [...otherScores, myScore];
  allScores.sort((a, b) => b - a);
  const rank = allScores.indexOf(myScore) + 1;

  if (rank === 1) return "Gold";
  if (rank === 2) return "Silver";
  if (rank === 3) return "Bronze";
  return "No Medal";
}
```
