---
id: 697a49e6ff50d756c9b6935e
title: "チャレンジ 181: 2026年冬季大会 2日目: スノーボード"
challengeType: 28
dashedName: challenge-181
---

# --description--

スノーボーダーの初期のスタンスと回転角度が与えられたとき、着地時のスタンスを判定してください。

- スノーボーダーのスタンスは `"Regular"` または `"Goofy"` のいずれかです。
- トリックの回転は90度の倍数です。正の値は時計回りの回転、負の値は反時計回りの回転を示します。
- 着地時のスタンスは180度回転するごとに反転します。

例えば、`"Regular"` と `90` が与えられたら `"Regular"` を返してください。`"Regular"` と `180` 度が与えられたら `"Goofy"` を返してください。

# --hints--

`getLandingStance("Regular", 90)` は `"Regular"` を返すべきです。

```js
assert.equal(getLandingStance("Regular", 90), "Regular");
```

`getLandingStance("Regular", 180)` は `"Goofy"` を返すべきです。

```js
assert.equal(getLandingStance("Regular", 180), "Goofy");
```

`getLandingStance("Goofy", -270)` は `"Regular"` を返すべきです。

```js
assert.equal(getLandingStance("Goofy", -270), "Regular");
```

`getLandingStance("Regular", 2340)` は `"Goofy"` を返すべきです。

```js
assert.equal(getLandingStance("Regular", 2340), "Goofy");
```

`getLandingStance("Goofy", 2160)` は `"Goofy"` を返すべきです。

```js
assert.equal(getLandingStance("Goofy", 2160), "Goofy");
```

`getLandingStance("Goofy", -540)` は `"Regular"` を返すべきです。

```js
assert.equal(getLandingStance("Goofy", -540), "Regular");
```

# --seed--

## --seed-contents--

```js
function getLandingStance(startStance, rotation) {

  return startStance;
}
```

# --solutions--

```js
function getLandingStance(startStance, rotation) {
  const flips = Math.floor(Math.abs(rotation) / 180);
  if (flips % 2 === 0) {
    return startStance;
  } else {
    return startStance === "Regular" ? "Goofy" : "Regular";
  }
}
```
