---
id: 697a49e6ff50d756c9b69364
title: "チャレンジ 187: 2026年冬季大会 第8日目: リュージュ"
challengeType: 28
dashedName: challenge-187
---

# --description--

5つの数字からなる配列が与えられます。各数字はリュージュ選手がトラックの区間を完走するのにかかった時間（秒）を表しています。どの区間の速度が最も速かったか、その速度がいくらだったかを判定してください。

トラックは以下の区間に分かれています：

- 区間1：320メートル
- 区間2：280メートル
- 区間3：350メートル
- 区間4：300メートル
- 区間5：250メートル

与えられた配列の最初の値は区間1の時間、2番目の値は区間2の時間、というように対応しています。

区間の速度（メートル毎秒）を計算するには、距離を時間で割ってください。

`"The luger's fastest speed was X m/s on segment Y."`を返してください。ここで`X`は小数点以下2桁に丸めた最速の速度、`Y`は最速速度が出た区間番号です。

# --hints--

`getFastestSpeed([9.523, 8.234, 10.012, 9.001, 7.128])`は`"The luger's fastest speed was 35.07 m/s on segment 5."`を返すべきです

```js
assert.equal(getFastestSpeed([9.523, 8.234, 10.012, 9.001, 7.128]), "The luger's fastest speed was 35.07 m/s on segment 5.");
```

`getFastestSpeed([9.381, 7.417, 9.912, 8.815, 7.284])`は`"The luger's fastest speed was 37.75 m/s on segment 2."`を返すべきです

```js
assert.equal(getFastestSpeed([9.381, 7.417, 9.912, 8.815, 7.284]), "The luger's fastest speed was 37.75 m/s on segment 2.");
```

`getFastestSpeed([8.890, 7.601, 9.093, 8.392, 6.912])`は`"The luger's fastest speed was 38.49 m/s on segment 3."`を返すべきです

```js
assert.equal(getFastestSpeed([8.890, 7.601, 9.093, 8.392, 6.912]), "The luger's fastest speed was 38.49 m/s on segment 3.");
```

`getFastestSpeed([8.490, 7.732, 10.103, 8.489, 6.840])`は`"The luger's fastest speed was 37.69 m/s on segment 1."`を返すべきです

```js
assert.equal(getFastestSpeed([8.490, 7.732, 10.103, 8.489, 6.840]), "The luger's fastest speed was 37.69 m/s on segment 1.");
```

`getFastestSpeed([8.204, 7.230, 9.673, 7.645, 6.508])`は`"The luger's fastest speed was 39.24 m/s on segment 4."`を返すべきです

```js
assert.equal(getFastestSpeed([8.204, 7.230, 9.673, 7.645, 6.508]), "The luger's fastest speed was 39.24 m/s on segment 4.");
```

# --seed--

## --seed-contents--

```js
function getFastestSpeed(times) {

  return times;
}
```

# --solutions--

```js
function getFastestSpeed(times) {
  const distances = [320, 280, 350, 300, 250];
  let fastestSpeed = 0;
  let fastestSegment = 0;

  for (let i = 0; i < times.length; i++) {
    const speed = distances[i] / times[i];
    if (speed > fastestSpeed) {
      fastestSpeed = speed;
      fastestSegment = i + 1;
    }
  }

  return `The luger's fastest speed was ${fastestSpeed.toFixed(2)} m/s on segment ${fastestSegment}.`;
}
```
