---
id: 697a49e9860d24853adef67b
title: "チャレンジ 190: 2026年冬季大会 11日目：アイスホッケー"
challengeType: 28
dashedName: challenge-190
---

# --description--

6つのアイスホッケーチームとラウンドロビンゲーム後の成績が与えられたとき、準決勝の対戦カードを決定してください。

- 各配列の項目はチームとその成績を `"TEAM: W-OTW-OTL-L"` の形式で持ちます。ここで：
  - `"W"` はレギュレーションでの勝利数で、1勝につき3ポイントです
  - `"OTW"` は延長戦での勝利数で、1勝につき2ポイントです
  - `"OTL"` は延長戦での敗北数で、1敗につき1ポイントです
  - `"L"` は敗北数で、ポイントは0です
  
例えば、`"FIN: 2-2-1-0"` は成績を合計して11ポイントになります。
  
各チームの合計ポイントを計算し、`"The semi-final games will be (1st) vs (4th) and (2nd) vs (3rd)."` を返してください。例えば、`"The semi-final games will be FIN vs SWE and CAN vs USA."`

# --hints--

`getSemifinalMatchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"])` は `"The semi-final games will be FIN vs SWE and CAN vs USA."` を返すべきです。

```js
assert.equal(getSemifinalMatchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"]), "The semi-final games will be FIN vs SWE and CAN vs USA.");
```

`getSemifinalMatchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"])` は `"The semi-final games will be USA vs CZE and CAN vs FIN."` を返すべきです。

```js
assert.equal(getSemifinalMatchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"]), "The semi-final games will be USA vs CZE and CAN vs FIN.");
```

`getSemifinalMatchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"])` は `"The semi-final games will be CAN vs ITA and USA vs CZE."` を返すべきです。

```js
assert.equal(getSemifinalMatchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"]), "The semi-final games will be CAN vs ITA and USA vs CZE.");
```

`getSemifinalMatchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"])` は `"The semi-final games will be JPN vs ITA and AUT vs KOR."` を返すべきです。

```js
assert.equal(getSemifinalMatchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"]), "The semi-final games will be JPN vs ITA and AUT vs KOR.");
```

# --seed--

## --seed-contents--

```js
function getSemifinalMatchups(teams) {

  return teams;
}
```

# --solutions--

```js
function getSemifinalMatchups(teams) {
    const teamPoints = teams.map(teamStr => {
        const [name, record] = teamStr.split(": ");
        const [W, OTW, OTL, L] = record.split("-").map(Number);
        const points = W * 3 + OTW * 2 + OTL * 1 + L * 0;
        return { name, points };
    });

    teamPoints.sort((a, b) => b.points - a.points);
    const [first, second, third, fourth] = teamPoints;

    return `The semi-final games will be ${first.name} vs ${fourth.name} and ${second.name} vs ${third.name}.`;
}
```
