---
id: 697a49e9860d24853adef67f
title: "チャレンジ 194: 2026年冬季大会 15日目: フリースタイルスキー"
challengeType: 28
dashedName: challenge-194
---

# --description--

2語からなるトリック名が、フリースタイルスキーの有効なトリック名かどうか判定してください。

トリックが有効なのは、最初の単語が有効な最初の単語のリストにあり、2番目の単語が有効な2番目の単語のリストにある場合です。

- 2語は1つのスペースで区切られます。

有効な最初の単語:

|`"Misty"`|
|-|
|`"Ghost"`|
|`"Thunder"`|
|`"Solar"`|
|`"Sky"`|
|`"Phantom"`|
|`"Frozen"`|
|`"Polar"`|

有効な2番目の単語:

|`"Twister"`|
|-|
|`"Icequake"`|
|`"Avalanche"`|
|`"Vortex"`|
|`"Snowstorm"`|
|`"Frostbite"`|
|`"Blizzard"`|
|`"Shadow"`|

# --hints--

`isValidTrick("Polar Vortex")` は `true` を返すべきです。

```js
assert.isTrue(isValidTrick("Polar Vortex"));
```

`isValidTrick("Solar Icequake")` は `true` を返すべきです。

```js
assert.isTrue(isValidTrick("Solar Icequake"));
```

`isValidTrick("Thunder Blizzard")` は `true` を返すべきです。

```js
assert.isTrue(isValidTrick("Thunder Blizzard"));
```

`isValidTrick("Phantom Frostbite")` は `true` を返すべきです。

```js
assert.isTrue(isValidTrick("Phantom Frostbite"));
```

`isValidTrick("Ghost Avalanche")` は `true` を返すべきです。

```js
assert.isTrue(isValidTrick("Ghost Avalanche"));
```

`isValidTrick("Snowstorm Shadow")` は `false` を返すべきです。

```js
assert.isFalse(isValidTrick("Snowstorm Shadow"));
```

`isValidTrick("Solar Sky")` は `false` を返すべきです。

```js
assert.isFalse(isValidTrick("Solar Sky"));
```

# --seed--

## --seed-contents--

```js
function isValidTrick(trickName) {

  return trickName;
}
```

# --solutions--

```js
function isValidTrick(trickName) {
  const validFirst = ["Misty", "Ghost", "Thunder", "Solar", "Sky", "Phantom", "Frozen", "Polar"];
  const validSecond = ["Twister", "Icequake", "Avalanche", "Vortex", "Snowstorm", "Frostbite", "Blizzard", "Shadow"];

  const words = trickName.split(" ");
  const [first, second] = words;
  return validFirst.includes(first) && validSecond.includes(second);
}
```
