---
id: 697a49e6ff50d756c9b6935d
title: "チャレンジ 180: 2026年冬季大会 第1日目：開会式"
challengeType: 28
dashedName: challenge-180
---

# --description--

今日は2026年冬季大会の開幕日です。これからの17日間は、この大会にちなんだコーディングチャレンジが続きます。

最初のチャレンジでは、2文字の国コードが与えられ、その国の旗の絵文字を返す必要があります。

以下のリストを使ってください：

| 国名                       | コード       | 旗       |
| -                         | -:          | :--:    |
| Albania                    |            `"AL"` | `"🇦🇱"` |
| Andorra                    |            `"AD"` | `"🇦🇩"` |
| Argentina                  |            `"AR"` | `"🇦🇷"` |
| Armenia                    |            `"AM"` | `"🇦🇲"` |
| Australia                  |            `"AU"` | `"🇦🇺"` |
| Austria                    |            `"AT"` | `"🇦🇹"` |
| Azerbaijan                 |            `"AZ"` | `"🇦🇿"` |
| Belgium                    |            `"BE"` | `"🇧🇪"` |
| Benin                      |            `"BJ"` | `"🇧🇯"` |
| Bolivia                    |            `"BO"` | `"🇧🇴"` |
| Bosnia and Herzegovina     |            `"BA"` | `"🇧🇦"` |
| Brazil                     |            `"BR"` | `"🇧🇷"` |
| Bulgaria                   |            `"BG"` | `"🇧🇬"` |
| Canada                     |            `"CA"` | `"🇨🇦"` |
| Chile                      |            `"CL"` | `"🇨🇱"` |
| China                      |            `"CN"` | `"🇨🇳"` |
| Colombia                   |            `"CO"` | `"🇨🇴"` |
| Croatia                    |            `"HR"` | `"🇭🇷"` |
| Cyprus                     |            `"CY"` | `"🇨🇾"` |
| Czech Republic             |            `"CZ"` | `"🇨🇿"` |
| Denmark                    |            `"DK"` | `"🇩🇰"` |
| Ecuador                    |            `"EC"` | `"🇪🇨"` |
| Eritrea                    |            `"ER"` | `"🇪🇷"` |
| Estonia                    |            `"EE"` | `"🇪🇪"` |
| Finland                    |            `"FI"` | `"🇫🇮"` |
| France                     |            `"FR"` | `"🇫🇷"` |
| Georgia                    |            `"GE"` | `"🇬🇪"` |
| Germany                    |            `"DE"` | `"🇩🇪"` |
| Great Britain              |            `"GB"` | `"🇬🇧"` |
| Greece                     |            `"GR"` | `"🇬🇷"` |
| Guinea-Bissau              |            `"GW"` | `"🇬🇼"` |
| Haiti                      |            `"HT"` | `"🇭🇹"` |
| Hong Kong                  |            `"HK"` | `"🇭🇰"` |
| Hungary                    |            `"HU"` | `"🇭🇺"` |
| Iceland                    |            `"IS"` | `"🇮🇸"` |
| India                      |            `"IN"` | `"🇮🇳"` |
| Iran                       |            `"IR"` | `"🇮🇷"` |
| Ireland                    |            `"IE"` | `"🇮🇪"` |
| Israel                     |            `"IL"` | `"🇮🇱"` |
| Italy                      |            `"IT"` | `"🇮🇹"` |
| Jamaica                    |            `"JM"` | `"🇯🇲"` |
| Japan                      |            `"JP"` | `"🇯🇵"` |
| Kazakhstan                 |            `"KZ"` | `"🇰🇿"` |
| Kenya                      |            `"KE"` | `"🇰🇪"` |
| Kosovo                     |            `"XK"` | `"🇽🇰"` |
| Kyrgyzstan                 |            `"KG"` | `"🇰🇬"` |
| Latvia                     |            `"LV"` | `"🇱🇻"` |
| Lebanon                    |            `"LB"` | `"🇱🇧"` |
| Liechtenstein              |            `"LI"` | `"🇱🇮"` |
| Lithuania                  |            `"LT"` | `"🇱🇹"` |
| Luxembourg                 |            `"LU"` | `"🇱🇺"` |
| Madagascar                 |            `"MG"` | `"🇲🇬"` |
| Malaysia                   |            `"MY"` | `"🇲🇾"` |
| Malta                      |            `"MT"` | `"🇲🇹"` |
| Mexico                     |            `"MX"` | `"🇲🇽"` |
| Moldova                    |            `"MD"` | `"🇲🇩"` |
| Monaco                     |            `"MC"` | `"🇲🇨"` |
| Mongolia                   |            `"MN"` | `"🇲🇳"` |
| Montenegro                 |            `"ME"` | `"🇲🇪"` |
| Morocco                    |            `"MA"` | `"🇲🇦"` |
| Netherlands                |            `"NL"` | `"🇳🇱"` |
| New Zealand                |            `"NZ"` | `"🇳🇿"` |
| Nigeria                    |            `"NG"` | `"🇳🇬"` |
| North Macedonia            |            `"MK"` | `"🇲🇰"` |
| Norway                     |            `"NO"` | `"🇳🇴"` |
| Pakistan                   |            `"PK"` | `"🇵🇰"` |
| Philippines                |            `"PH"` | `"🇵🇭"` |
| Poland                     |            `"PL"` | `"🇵🇱"` |
| Portugal                   |            `"PT"` | `"🇵🇹"` |
| Puerto Rico                |            `"PR"` | `"🇵🇷"` |
| Romania                    |            `"RO"` | `"🇷🇴"` |
| San Marino                 |            `"SM"` | `"🇸🇲"` |
| Saudi Arabia               |            `"SA"` | `"🇸🇦"` |
| Serbia                     |            `"RS"` | `"🇷🇸"` |
| Singapore                  |            `"SG"` | `"🇸🇬"` |
| Slovakia                   |            `"SK"` | `"🇸🇰"` |
| Slovenia                   |            `"SI"` | `"🇸🇮"` |
| South Africa               |            `"ZA"` | `"🇿🇦"` |
| South Korea                |            `"KR"` | `"🇰🇷"` |
| Spain                      |            `"ES"` | `"🇪🇸"` |
| Sweden                     |            `"SE"` | `"🇸🇪"` |
| Switzerland                |            `"CH"` | `"🇨🇭"` |
| Thailand                   |            `"TH"` | `"🇹🇭"` |
| Trinidad & Tobago          |            `"TT"` | `"🇹🇹"` |
| Turkey                     |            `"TR"` | `"🇹🇷"` |
| Ukraine                    |            `"UA"` | `"🇺🇦"` |
| United Arab Emirates       |            `"AE"` | `"🇦🇪"` |
| United States              |            `"US"` | `"🇺🇸"` |
| Uruguay                    |            `"UY"` | `"🇺🇾"` |
| Uzbekistan                 |            `"UZ"` | `"🇺🇿"` |
| Venezuela                  |            `"VE"` | `"🇻🇪"` |

# --hints--

`getFlag("AL")` は `"🇦🇱"` を返す必要があります。

```js
assert.equal(getFlag("AL"), "🇦🇱");
```

`getFlag("AD")` は `"🇦🇩"` を返す必要があります。

```js
assert.equal(getFlag("AD"), "🇦🇩");
```

`getFlag("AR")` は `"🇦🇷"` を返す必要があります。

```js
assert.equal(getFlag("AR"), "🇦🇷")
```

`getFlag("AM")` は `"🇦🇲"` を返す必要があります。

```js
assert.equal(getFlag("AM"), "🇦🇲")
```

`getFlag("AU")` は `"🇦🇺"` を返す必要があります。

```js
assert.equal(getFlag("AU"), "🇦🇺")
```

`getFlag("AT")` は `"🇦🇹"` を返す必要があります。

```js
assert.equal(getFlag("AT"), "🇦🇹")
```

`getFlag("AZ")` は `"🇦🇿"` を返す必要があります。

```js
assert.equal(getFlag("AZ"), "🇦🇿")
```

`getFlag("BE")` は `"🇧🇪"` を返す必要があります。

```js
assert.equal(getFlag("BE"), "🇧🇪")
```

`getFlag("BJ")` は `"🇧🇯"` を返す必要があります。

```js
assert.equal(getFlag("BJ"), "🇧🇯")
```

`getFlag("BO")` は `"🇧🇴"` を返す必要があります。

```js
assert.equal(getFlag("BO"), "🇧🇴")
```

`getFlag("BA")` は `"🇧🇦"` を返す必要があります。

```js
assert.equal(getFlag("BA"), "🇧🇦")
```

`getFlag("BR")` は `"🇧🇷"` を返す必要があります。

```js
assert.equal(getFlag("BR"), "🇧🇷")
```

`getFlag("BG")` は `"🇧🇬"` を返す必要があります。

```js
assert.equal(getFlag("BG"), "🇧🇬")
```

`getFlag("CA")` は `"🇨🇦"` を返す必要があります。

```js
assert.equal(getFlag("CA"), "🇨🇦")
```

`getFlag("CL")` は `"🇨🇱"` を返す必要があります。

```js
assert.equal(getFlag("CL"), "🇨🇱")
```

`getFlag("CN")` は `"🇨🇳"` を返す必要があります。

```js
assert.equal(getFlag("CN"), "🇨🇳")
```

`getFlag("CO")` は `"🇨🇴"` を返す必要があります。

```js
assert.equal(getFlag("CO"), "🇨🇴")
```

`getFlag("HR")` は `"🇭🇷"` を返す必要があります。

```js
assert.equal(getFlag("HR"), "🇭🇷")
```

`getFlag("CY")` は `"🇨🇾"` を返す必要があります。

```js
assert.equal(getFlag("CY"), "🇨🇾")
```

`getFlag("CZ")` は `"🇨🇿"` を返す必要があります。

```js
assert.equal(getFlag("CZ"), "🇨🇿")
```

`getFlag("DK")` は `"🇩🇰"` を返す必要があります。

```js
assert.equal(getFlag("DK"), "🇩🇰")
```

`getFlag("EC")` は `"🇪🇨"` を返す必要があります。

```js
assert.equal(getFlag("EC"), "🇪🇨")
```

`getFlag("ER")` は `"🇪🇷"` を返す必要があります。

```js
assert.equal(getFlag("ER"), "🇪🇷")
```

`getFlag("EE")` は `"🇪🇪"` を返す必要があります。

```js
assert.equal(getFlag("EE"), "🇪🇪")
```

`getFlag("FI")` は `"🇫🇮"` を返す必要があります。

```js
assert.equal(getFlag("FI"), "🇫🇮")
```

`getFlag("FR")` は `"🇫🇷"` を返す必要があります。

```js
assert.equal(getFlag("FR"), "🇫🇷")
```

`getFlag("GE")` は `"🇬🇪"` を返す必要があります。

```js
assert.equal(getFlag("GE"), "🇬🇪")
```

`getFlag("DE")` は `"🇩🇪"` を返す必要があります。

```js
assert.equal(getFlag("DE"), "🇩🇪")
```

`getFlag("GB")` は `"🇬🇧"` を返す必要があります。

```js
assert.equal(getFlag("GB"), "🇬🇧")
```

`getFlag("GR")` は `"🇬🇷"` を返す必要があります。

```js
assert.equal(getFlag("GR"), "🇬🇷")
```

`getFlag("GW")` は `"🇬🇼"` を返す必要があります。

```js
assert.equal(getFlag("GW"), "🇬🇼")
```

`getFlag("HT")` は `"🇭🇹"` を返す必要があります。

```js
assert.equal(getFlag("HT"), "🇭🇹")
```

`getFlag("HK")` は `"🇭🇰"` を返す必要があります。

```js
assert.equal(getFlag("HK"), "🇭🇰")
```

`getFlag("HU")` は `"🇭🇺"` を返す必要があります。

```js
assert.equal(getFlag("HU"), "🇭🇺")
```

`getFlag("IS")` は `"🇮🇸"` を返す必要があります。

```js
assert.equal(getFlag("IS"), "🇮🇸")
```

`getFlag("IN")` は `"🇮🇳"` を返す必要があります。

```js
assert.equal(getFlag("IN"), "🇮🇳")
```

`getFlag("IR")` は `"🇮🇷"` を返す必要があります。

```js
assert.equal(getFlag("IR"), "🇮🇷")
```

`getFlag("IE")` は `"🇮🇪"` を返す必要があります。

```js
assert.equal(getFlag("IE"), "🇮🇪")
```

`getFlag("IL")` は `"🇮🇱"` を返す必要があります。

```js
assert.equal(getFlag("IL"), "🇮🇱")
```

`getFlag("IT")` は `"🇮🇹"` を返す必要があります。

```js
assert.equal(getFlag("IT"), "🇮🇹")
```

`getFlag("JM")` は `"🇯🇲"` を返す必要があります。

```js
assert.equal(getFlag("JM"), "🇯🇲")
```

`getFlag("JP")` は `"🇯🇵"` を返す必要があります。

```js
assert.equal(getFlag("JP"), "🇯🇵")
```

`getFlag("KZ")` は `"🇰🇿"` を返す必要があります。

```js
assert.equal(getFlag("KZ"), "🇰🇿")
```

`getFlag("KE")` は `"🇰🇪"` を返す必要があります。

```js
assert.equal(getFlag("KE"), "🇰🇪")
```

`getFlag("XK")` は `"🇽🇰"` を返す必要があります。

```js
assert.equal(getFlag("XK"), "🇽🇰")
```

`getFlag("KG")` は `"🇰🇬"` を返す必要があります。

```js
assert.equal(getFlag("KG"), "🇰🇬")
```

`getFlag("LV")` は `"🇱🇻"` を返す必要があります。

```js
assert.equal(getFlag("LV"), "🇱🇻")
```

`getFlag("LB")` は `"🇱🇧"` を返す必要があります。

```js
assert.equal(getFlag("LB"), "🇱🇧")
```

`getFlag("LI")` は `"🇱🇮"` を返す必要があります。

```js
assert.equal(getFlag("LI"), "🇱🇮")
```

`getFlag("LT")` は `"🇱🇹"` を返す必要があります。

```js
assert.equal(getFlag("LT"), "🇱🇹")
```

`getFlag("LU")` は `"🇱🇺"` を返す必要があります。

```js
assert.equal(getFlag("LU"), "🇱🇺")
```

`getFlag("MG")` は `"🇲🇬"` を返す必要があります。

```js
assert.equal(getFlag("MG"), "🇲🇬")
```

`getFlag("MY")` は `"🇲🇾"` を返す必要があります。

```js
assert.equal(getFlag("MY"), "🇲🇾")
```

`getFlag("MT")` は `"🇲🇹"` を返す必要があります。

```js
assert.equal(getFlag("MT"), "🇲🇹")
```

`getFlag("MX")` は `"🇲🇽"` を返す必要があります。

```js
assert.equal(getFlag("MX"), "🇲🇽")
```

`getFlag("MD")` は `"🇲🇩"` を返す必要があります。

```js
assert.equal(getFlag("MD"), "🇲🇩")
```

`getFlag("MC")` は `"🇲🇨"` を返す必要があります。

```js
assert.equal(getFlag("MC"), "🇲🇨")
```

`getFlag("MN")` は `"🇲🇳"` を返す必要があります。

```js
assert.equal(getFlag("MN"), "🇲🇳")
```

`getFlag("ME")` は `"🇲🇪"` を返す必要があります。

```js
assert.equal(getFlag("ME"), "🇲🇪")
```

`getFlag("MA")` は `"🇲🇦"` を返す必要があります。

```js
assert.equal(getFlag("MA"), "🇲🇦")
```

`getFlag("NL")` は `"🇳🇱"` を返す必要があります。

```js
assert.equal(getFlag("NL"), "🇳🇱")
```

`getFlag("NZ")` は `"🇳🇿"` を返す必要があります。

```js
assert.equal(getFlag("NZ"), "🇳🇿")
```

`getFlag("NG")` は `"🇳🇬"` を返す必要があります。

```js
assert.equal(getFlag("NG"), "🇳🇬")
```

`getFlag("MK")` は `"🇲🇰"` を返す必要があります。

```js
assert.equal(getFlag("MK"), "🇲🇰")
```

`getFlag("NO")` は `"🇳🇴"` を返す必要があります。

```js
assert.equal(getFlag("NO"), "🇳🇴")
```

`getFlag("PK")` は `"🇵🇰"` を返す必要があります。

```js
assert.equal(getFlag("PK"), "🇵🇰")
```

`getFlag("PH")` は `"🇵🇭"` を返す必要があります。

```js
assert.equal(getFlag("PH"), "🇵🇭")
```

`getFlag("PL")` は `"🇵🇱"` を返す必要があります。

```js
assert.equal(getFlag("PL"), "🇵🇱")
```

`getFlag("PT")` は `"🇵🇹"` を返す必要があります。

```js
assert.equal(getFlag("PT"), "🇵🇹")
```

`getFlag("PR")` は `"🇵🇷"` を返す必要があります。

```js
assert.equal(getFlag("PR"), "🇵🇷")
```

`getFlag("RO")` は `"🇷🇴"` を返す必要があります。

```js
assert.equal(getFlag("RO"), "🇷🇴")
```

`getFlag("SM")` は `"🇸🇲"` を返す必要があります。

```js
assert.equal(getFlag("SM"), "🇸🇲")
```

`getFlag("SA")` は `"🇸🇦"` を返す必要があります。

```js
assert.equal(getFlag("SA"), "🇸🇦")
```

`getFlag("RS")` は `"🇷🇸"` を返す必要があります。

```js
assert.equal(getFlag("RS"), "🇷🇸")
```

`getFlag("SG")` は `"🇸🇬"` を返す必要があります。

```js
assert.equal(getFlag("SG"), "🇸🇬")
```

`getFlag("SK")` は `"🇸🇰"` を返す必要があります。

```js
assert.equal(getFlag("SK"), "🇸🇰")
```

`getFlag("SI")` は `"🇸🇮"` を返す必要があります。

```js
assert.equal(getFlag("SI"), "🇸🇮")
```

`getFlag("ZA")` は `"🇿🇦"` を返す必要があります。

```js
assert.equal(getFlag("ZA"), "🇿🇦")
```

`getFlag("KR")` は `"🇰🇷"` を返す必要があります。

```js
assert.equal(getFlag("KR"), "🇰🇷")
```

`getFlag("ES")` は `"🇪🇸"` を返す必要があります。

```js
assert.equal(getFlag("ES"), "🇪🇸")
```

`getFlag("SE")` は `"🇸🇪"` を返す必要があります。

```js
assert.equal(getFlag("SE"), "🇸🇪")
```

`getFlag("CH")` は `"🇨🇭"` を返す必要があります。

```js
assert.equal(getFlag("CH"), "🇨🇭")
```

`getFlag("TH")` は `"🇹🇭"` を返す必要があります。

```js
assert.equal(getFlag("TH"), "🇹🇭")
```

`getFlag("TT")` は `"🇹🇹"` を返す必要があります。

```js
assert.equal(getFlag("TT"), "🇹🇹")
```

`getFlag("TR")` は `"🇹🇷"` を返す必要があります。

```js
assert.equal(getFlag("TR"), "🇹🇷")
```

`getFlag("UA")` は `"🇺🇦"` を返す必要があります。

```js
assert.equal(getFlag("UA"), "🇺🇦")
```

`getFlag("AE")` は `"🇦🇪"` を返す必要があります。

```js
assert.equal(getFlag("AE"), "🇦🇪")
```

`getFlag("US")` は `"🇺🇸"` を返す必要があります。

```js
assert.equal(getFlag("US"), "🇺🇸")
```

`getFlag("UY")` は `"🇺🇾"` を返す必要があります。

```js
assert.equal(getFlag("UY"), "🇺🇾")
```

`getFlag("UZ")` は `"🇺🇿"` を返す必要があります。

```js
assert.equal(getFlag("UZ"), "🇺🇿")
```

`getFlag("VE")` は `"🇻🇪"` を返す必要があります。

```js
assert.equal(getFlag("VE"), "🇻🇪")
```

# --seed--

## --seed-contents--

```js
function getFlag(code) {

  return code;
}
```

# --solutions--

```js
function getFlag(code) {
  return [...code]
    .map(char => String.fromCodePoint(127397 + char.charCodeAt()))
    .join("");
}
```
