---
id: 697a49e9860d24853adef67e
title: "チャレンジ 193: 2026年冬季大会 14日目: スキーマウンテニアリング"
challengeType: 28
dashedName: challenge-193
---

# --description--

山の積雪深と斜面の傾斜から、雪崩のリスクがあるかどうかを判定してください。

- 積雪深の値は `"Shallow"`、`"Moderate"`、または `"Deep"` です。
- 傾斜の値は `"Gentle"`、`"Steep"`、または `"Very Steep"` です。

このテーブルに基づいて `"Safe"` または `"Risky"` を返してください：

| |`"Shallow"`|`"Moderate"`|`"Deep"`|
|-|-|-|-|
|`"Gentle"`|`"Safe"`|`"Safe"`|`"Safe"`|
|`"Steep"`|`"Safe"`|`"Risky"`|`"Risky"`|
|`"Very Steep"`|`"Safe"`| `"Risky"`|`"Risky"`|

# --hints--

`avalancheRisk("Shallow", "Gentle")` は `"Safe"` を返すべきです。

```js
assert.equal(avalancheRisk("Shallow", "Gentle"), "Safe");
```

`avalancheRisk("Shallow", "Steep")` は `"Safe"` を返すべきです。

```js
assert.equal(avalancheRisk("Shallow", "Steep"), "Safe");
```

`avalancheRisk("Shallow", "Very Steep")` は `"Safe"` を返すべきです。

```js
assert.equal(avalancheRisk("Shallow", "Very Steep"), "Safe");
```

`avalancheRisk("Moderate", "Gentle")` は `"Safe"` を返すべきです。

```js
assert.equal(avalancheRisk("Moderate", "Gentle"), "Safe");
```

`avalancheRisk("Moderate", "Steep")` は `"Risky"` を返すべきです。

```js
assert.equal(avalancheRisk("Moderate", "Steep"), "Risky");
```

`avalancheRisk("Moderate", "Very Steep")` は `"Risky"` を返すべきです。

```js
assert.equal(avalancheRisk("Moderate", "Very Steep"), "Risky");
```

`avalancheRisk("Deep", "Gentle")` は `"Safe"` を返すべきです。

```js
assert.equal(avalancheRisk("Deep", "Gentle"), "Safe");
```

`avalancheRisk("Deep", "Steep")` は `"Risky"` を返すべきです。

```js
assert.equal(avalancheRisk("Deep", "Steep"), "Risky");
```

`avalancheRisk("Deep", "Very Steep")` は `"Risky"` を返すべきです。

```js
assert.equal(avalancheRisk("Deep", "Very Steep"), "Risky");
```

# --seed--

## --seed-contents--

```js
function avalancheRisk(snowDepth, slope) {

  return snowDepth;
}
```

# --solutions--

```js
function avalancheRisk(snowDepth, slope) {
  const riskTable = {
    Gentle: {
      Shallow: "Safe",
      Moderate: "Safe",
      Deep: "Safe",
    },
    Steep: {
      Shallow: "Safe",
      Moderate: "Risky",
      Deep: "Risky",
    },
    "Very Steep": {
      Shallow: "Safe",
      Moderate: "Risky",
      Deep: "Risky",
    },
  };

  return riskTable[slope][snowDepth];
}
```
