---
id: 697a49e9860d24853adef67c
title: "チャレンジ 191: 2026年冬季大会 12日目: ボブスレー"
challengeType: 28
dashedName: challenge-191
---

# --description--

ボブスレーチームの選手の体重を表す配列とボブスレー自体の重さを表す数値が与えられたとき、チームがレースに出場できるかどうかを判定してください。

- 配列の長さはチームの人数を決めます：1人、2人、または4人チームです。
- すべての体重の値はキログラム（kg）単位です。

ボブスレー（そり単体）の最小重量は以下の通りです：

- 1人チームの場合は162kg
- 2人チームの場合は170kg
- 4人チームの場合は210kg

ボブスレーの総重量（選手の体重＋そりの重さ）は以下を超えてはいけません：

- 1人チームの場合は247kg
- 2人チームの場合は390kg
- 4人チームの場合は630kg

チームがすべての条件を満たす場合は`"Eligible"`を、1つ以上の条件を満たさない場合は`"Not Eligible"`を返してください。

# --hints--

`checkEligibility([78], 165)`は`"Eligible"`を返すべきです。

```js
assert.equal(checkEligibility([78], 165), "Eligible");
```

`checkEligibility([80], 160)`は`"Not Eligible"`を返すべきです。

```js
assert.equal(checkEligibility([80], 160), "Not Eligible");
```

`checkEligibility([80], 170)`は`"Not Eligible"`を返すべきです。

```js
assert.equal(checkEligibility([80], 170), "Not Eligible");
```

`checkEligibility([85, 90], 170)`は`"Eligible"`を返すべきです。

```js
assert.equal(checkEligibility([85, 90], 170), "Eligible");
```

`checkEligibility([85, 95], 168)`は`"Not Eligible"`を返すべきです。

```js
assert.equal(checkEligibility([85, 95], 168), "Not Eligible");
```

`checkEligibility([112, 97], 185)`は`"Not Eligible"`を返すべきです。

```js
assert.equal(checkEligibility([112, 97], 185), "Not Eligible");
```

`checkEligibility([110, 102, 90, 106], 222)`は`"Eligible"`を返すべきです。

```js
assert.equal(checkEligibility([110, 102, 90, 106], 222), "Eligible");
```

`checkEligibility([106, 99, 90, 88], 205)`は`"Not Eligible"`を返すべきです。

```js
assert.equal(checkEligibility([106, 99, 90, 88], 205), "Not Eligible");
```

`checkEligibility([106, 99, 103, 96], 227)`は`"Not Eligible"`を返すべきです。

```js
assert.equal(checkEligibility([106, 99, 103, 96], 227), "Not Eligible");
```

# --seed--

## --seed-contents--

```js
function checkEligibility(athleteWeights, sledWeight) {

  return athleteWeights;
}
```

# --solutions--

```js
function checkEligibility(athleteWeights, sledWeight) {
  const teamSize = athleteWeights.length;

  const rules = {
    1: { minSled: 162, maxTotal: 247 },
    2: { minSled: 170, maxTotal: 390 },
    4: { minSled: 210, maxTotal: 630 }
  };

  const { minSled, maxTotal } = rules[teamSize];
  const athleteTotal = athleteWeights.reduce((sum, w) => sum + w, 0);
  const totalWeight = athleteTotal + sledWeight;

  if (sledWeight < minSled) {
    return "Not Eligible";
  }

  if (totalWeight > maxTotal) {
    return "Not Eligible";
  }

  return "Eligible";
}
```
