---
id: 697a49e6ff50d756c9b6935f
title: "チャレンジ 182: 2026年冬季大会3日目：バイアスロン"
challengeType: 28
dashedName: challenge-182
---

# --description--

整数の配列が与えられます。各値はバイアスロンの1ラウンドで命中したターゲットの数を表しています。選手がスキーで走らなければならない合計のペナルティ距離を返してください。

- 各ラウンドは5つのターゲットで構成されています。
- 外したターゲット1つにつき150メートルのペナルティループがあります。

# --hints--

`calculatePenaltyDistance([4, 4])` は `300` を返すべきです。

```js
assert.equal(calculatePenaltyDistance([4, 4]), 300);
```

`calculatePenaltyDistance([5, 5])` は `0` を返すべきです。

```js
assert.equal(calculatePenaltyDistance([5, 5]), 0);
```

`calculatePenaltyDistance([4, 5, 3, 5])` は `450` を返すべきです。

```js
assert.equal(calculatePenaltyDistance([4, 5, 3, 5]), 450);
```

`calculatePenaltyDistance([5, 4, 5, 5])` は `150` を返すべきです。

```js
assert.equal(calculatePenaltyDistance([5, 4, 5, 5]), 150);
```

`calculatePenaltyDistance([4, 3, 0, 3])` は `1500` を返すべきです。

```js
assert.equal(calculatePenaltyDistance([4, 3, 0, 3]), 1500);
```

# --seed--

## --seed-contents--

```js
function calculatePenaltyDistance(rounds) {

  return rounds;
}
```

# --solutions--

```js
function calculatePenaltyDistance(rounds) {
  let totalPenalty = 0;

  for (const hits of rounds) {
    const misses = 5 - hits;
    totalPenalty += misses * 150;
  }

  return totalPenalty;
}
```
