---
id: 697a49e6ff50d756c9b69366
title: "チャレンジ 189: 2026年冬季大会 10日目：アルペンスキー"
challengeType: 28
dashedName: challenge-189
---

# --description--

スキー場の垂直落差、水平距離、タイプが与えられたとき、その斜面の難易度評価を判定してください。

評価を決めるには：

- 落差を距離で割って斜面の傾斜を計算します。
- 次に、斜面のタイプに基づいて調整された傾斜を計算します：
  - `"Downhill"`: 傾斜に1.2を乗算する
  - `"Slalom"`: 傾斜に0.9を乗算する
  - `"Giant Slalom"`: 傾斜に1.0を乗算する

返す値：

- 調整された傾斜が0.1以下の場合は`"Green"`
- 調整された傾斜が0.1より大きく0.25以下の場合は`"Blue"`
- 調整された傾斜が0.25より大きい場合は`"Black"`

# --hints--

`getHillRating(95, 900, "Slalom")`は`"Green"`を返すべきです。

```js
assert.equal(getHillRating(95, 900, "Slalom"), "Green");
```

`getHillRating(620, 2800, "Downhill")`は`"Black"`を返すべきです。

```js
assert.equal(getHillRating(620, 2800, "Downhill"), "Black");
```

`getHillRating(420, 1680, "Giant Slalom")`は`"Blue"`を返すべきです。

```js
assert.equal(getHillRating(420, 1680, "Giant Slalom"), "Blue");
```

`getHillRating(250, 3000, "Downhill")`は`"Green"`を返すべきです。

```js
assert.equal(getHillRating(250, 3000, "Downhill"), "Green");
```

`getHillRating(110, 900, "Slalom")`は`"Blue"`を返すべきです。

```js
assert.equal(getHillRating(110, 900, "Slalom"), "Blue");
```

`getHillRating(380, 1500, "Giant Slalom")`は`"Black"`を返すべきです。

```js
assert.equal(getHillRating(380, 1500, "Giant Slalom"), "Black");
```

# --seed--

## --seed-contents--

```js
function getHillRating(drop, distance, type) {

  return drop;
}
```

# --solutions--

```js
function getHillRating(drop, distance, type) {
  let steepness = drop / distance;

  if (type === "Downhill") {
    steepness *= 1.2;
  } else if (type === "Slalom") {
    steepness *= 0.9;
  } else if (type === "Giant Slalom") {
    steepness *= 1.0;
  }

  if (steepness <= 0.1) {
    return "Green";
  } else if (steepness <= 0.25) {
    return "Blue";
  } else {
    return "Black";
  }
}
```
