---
id: 697a49e6ff50d756c9b69361
title: "チャレンジ 184: 2026年冬季大会 第5日目: クロスカントリースキー"
challengeType: 28
dashedName: challenge-184
---

# --description--

クロスカントリースキーのレースの完走時間の配列が与えられたとき、それらを優勝者との差の時間に変換してください。

- 与えられる時間は`"H:MM:SS"`形式の文字列です。
- 与えられる時間は速い順から遅い順に並んでいます。
- 優勝者の時間（最速の時間）は`"0"`に対応します。
- 他の各時間は、`"+M:SS"`形式で優勝者との差の時間を示す必要があります。

例えば、`["1:25:32", "1:26:10", "1:27:05"]`が与えられた場合、`["0", "+0:38", "+1:33"]`を返してください。

# --hints--

`getRelativeResults(["1:25:32", "1:26:10", "1:27:05"])`は`["0", "+0:38", "+1:33"]`を返すべきです。

```js
assert.deepEqual(getRelativeResults(["1:25:32", "1:26:10", "1:27:05"]), ["0", "+0:38", "+1:33"]);
```

`getRelativeResults(["1:00:01", "1:00:05", "1:00:10"])`は`["0", "+0:04", "+0:09"]`を返すべきです。

```js
assert.deepEqual(getRelativeResults(["1:00:01", "1:00:05", "1:00:10"]), ["0", "+0:04", "+0:09"]);
```

`getRelativeResults(["1:10:06", "1:10:23", "1:10:48", "1:12:11"])`は`["0", "+0:17", "+0:42", "+2:05"]`を返すべきです。

```js
assert.deepEqual(getRelativeResults(["1:10:06", "1:10:23", "1:10:48", "1:12:11"]), ["0", "+0:17", "+0:42", "+2:05"]);
```

`getRelativeResults(["0:49:13", "0:49:15", "0:50:14", "0:51:30", "0:51:58", "0:52:16", "0:53:12", "0:53:31", "0:56:19", "1:02:20"])`は`["0", "+0:02", "+1:01", "+2:17", "+2:45", "+3:03", "+3:59", "+4:18", "+7:06", "+13:07"]`を返すべきです。

```js
assert.deepEqual(getRelativeResults(["0:49:13", "0:49:15", "0:50:14", "0:51:30", "0:51:58", "0:52:16", "0:53:12", "0:53:31", "0:56:19", "1:02:20"]), ["0", "+0:02", "+1:01", "+2:17", "+2:45", "+3:03", "+3:59", "+4:18", "+7:06", "+13:07"]);
```

`getRelativeResults(["2:01:15", "2:10:45", "2:10:53", "2:11:04", "2:11:55", "2:13:27", "2:14:30", "2:15:10"])`は`["0", "+9:30", "+9:38", "+9:49", "+10:40", "+12:12", "+13:15", "+13:55"]`を返すべきです。

```js
assert.deepEqual(getRelativeResults(["2:01:15", "2:10:45", "2:10:53", "2:11:04", "2:11:55", "2:13:27", "2:14:30", "2:15:10"]), ["0", "+9:30", "+9:38", "+9:49", "+10:40", "+12:12", "+13:15", "+13:55"]);
```

# --seed--

## --seed-contents--

```js
function getRelativeResults(results) {

  return results;
}
```

# --solutions--

```js
function getRelativeResults(results) {
  const timeToSeconds = (timeStr) => {
    const [hours, minutes, seconds] = timeStr.split(':').map(Number);
    return hours * 3600 + minutes * 60 + seconds;
  };

  const secondsToTimeFormat = (secs) => {
    const mins = Math.floor(secs / 60);
    const secsRemainder = secs % 60;
    return `+${mins}:${String(secsRemainder).padStart(2, '0')}`;
  };

  const winnerSeconds = timeToSeconds(results[0]);

  return results.map((time, index) => {
    if (index === 0) return '0';
    const currentSeconds = timeToSeconds(time);
    const difference = currentSeconds - winnerSeconds;
    return secondsToTimeFormat(difference);
  });
}
```
