---
id: 698a1a73ade5ac0e19180fa4
title: "チャレンジ 201: マトリックスシフト"
challengeType: 28
dashedName: challenge-201
---

# --description--

数値のマトリックス（配列の配列）と整数が与えられたとき、マトリックス内のすべての値を指定された量だけシフトしてください。

- 正のシフトは値を右に移動させます。
- 負のシフトは値を左に移動させます。

値はラップ（循環）します：

- マトリックスを1つの連続したシーケンスとして扱います。
- 値が行の終わりを超えて移動すると、次の行の先頭から続きます。
- 値がマトリックスの最後の位置を超えると、最初の位置に戻ります。
- 左にシフトするときも同様に逆方向で適用されます。

例えば、以下が与えられた場合：

```js
[
  [1, 2, 3],
  [4, 5, 6]
]
```

`1` のシフトで、すべての数字を右に1つ移動させます：

```js
[
  [6, 1, 2],
  [3, 4, 5]
]
```

# --hints--

`shiftMatrix([[1, 2, 3], [4, 5, 6]], 1)` は `[[6, 1, 2], [3, 4, 5]]` を返すべきです。

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6]], 1), [[6, 1, 2], [3, 4, 5]]);
```

`shiftMatrix([[1, 2, 3], [4, 5, 6]], -1)` は `[[2, 3, 4], [5, 6, 1]]` を返すべきです。

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6]], -1), [[2, 3, 4], [5, 6, 1]]);
```

`shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5)` は `[[5, 6, 7], [8, 9, 1], [2, 3, 4]]` を返すべきです。

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5), [[5, 6, 7], [8, 9, 1], [2, 3, 4]]);
```

`shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6)` は `[[7, 8, 9], [1, 2, 3], [4, 5, 6]]` を返すべきです。

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6), [[7, 8, 9], [1, 2, 3], [4, 5, 6]]);
```

`shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7)` は `[[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]]` を返すべきです。

```js
assert.deepEqual(shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7), [[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]]);
```

`shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54)` は `[[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]]` を返すべきです。

```js
assert.deepEqual(shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54), [[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]]);
```

# --seed--

## --seed-contents--

```js
function shiftMatrix(matrix, shift) {

  return matrix;
}
```

# --solutions--

```js
function shiftMatrix(matrix, shift) {
  const rows = matrix.length;
  const cols = matrix[0].length;
  const flat = matrix.flat();
  const n = flat.length;
  const s = ((shift % n) + n) % n;
  const shifted = flat.slice(-s).concat(flat.slice(0, n - s));
  const result = [];
  for (let i = 0; i < rows; i++) {
    result.push(shifted.slice(i * cols, (i + 1) * cols));
  }

  return result;
}
```
