---
id: 698a1a863194f1f4e63f645f
title: "チャレンジ 208: トレイルの移動"
challengeType: 28
dashedName: challenge-208
---

# --description--

トレイルマップを表す文字列の配列が与えられたとき、目的地に到達するために必要な移動の文字列を返してください。

与えられた文字列には以下の値が含まれます：

- `"C"`: 現在の位置
- `"G"`: 目的地
- `"T"`: 通行可能なトレイルの部分
- `"-"`: 通行不可能なマップの部分

現在地から目的地までトレイルに沿って進むために必要な移動を表す文字列を返してください。移動は以下の通りです：

- `"R"`: 右への移動
- `"D"`: 下への移動
- `"L"`: 左への移動
- `"U"`: 上への移動

- 現在地から目的地まで、分岐のない連続したトレイルが必ず1つだけ存在します。
- 各トレイルの位置は、最大で2つの通行可能な場所と接しています。

例えば、以下が与えられた場合：

```js
[
  "-CT--",
  "--T--",
  "--TT-",
  "---T-",
  "---G-"
]
```

`"RDDRDD"` を返してください。

# --hints--

`navigateTrail(["-CT--", "--T--", "--TT-", "---T-", "---G-"])` は `"RDDRDD"` を返すべきです。

```js
assert.equal(navigateTrail(["-CT--", "--T--", "--TT-", "---T-", "---G-"]), "RDDRDD");
```

`navigateTrail(["-----", "--TTG", "--T--", "--T--", "CTT--"])` は `"RRUUURR"` を返すべきです。

```js
assert.equal(navigateTrail(["-----", "--TTG", "--T--", "--T--", "CTT--"]), "RRUUURR");
```

`navigateTrail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"])` は `"DLDDRRRRD"` を返すべきです。

```js
assert.equal(navigateTrail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"]), "DLDDRRRRD");
```

`navigateTrail(["--------", "-CTTT---", "----T---", "---GT---", "--------"])` は `"RRRDDL"` を返すべきです。

```js
assert.equal(navigateTrail(["--------", "-CTTT---", "----T---", "---GT---", "--------"]), "RRRDDL");
```

`navigateTrail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"])` は `"ULLLUUURRRRRRDDDR"` を返すべきです。

```js
assert.equal(navigateTrail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"]), "ULLLUUURRRRRRDDDR");
```

# --seed--

## --seed-contents--

```js
function navigateTrail(map) {

  return map;
}
```

# --solutions--

```js
function navigateTrail(map) {
  const rows = map.length;
  const cols = map[0].length;

  const directions = [
    [0, 1, "R"],
    [1, 0, "D"],
    [0, -1, "L"],
    [-1, 0, "U"],
  ];

  let row, col;
  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (map[r][c] === "C") {
        row = r;
        col = c;
      }
    }
  }

  let result = "";
  let prev = null;

  while (map[row][col] !== "G") {
    for (let [dr, dc, move] of directions) {
      const newRow = row + dr;
      const newCol = col + dc;

      if (
        newRow >= 0 &&
        newRow < rows &&
        newCol >= 0 &&
        newCol < cols &&
        (map[newRow][newCol] === "T" ||
          map[newRow][newCol] === "G") &&
        (!prev || newRow !== prev[0] || newCol !== prev[1])
      ) {
        result += move;
        prev = [row, col];
        row = newRow;
        col = newCol;
        break;
      }
    }
  }

  return result;
}
```
