---
id: 698a1a73ade5ac0e19180fa0
title: "チャレンジ 197: 血液型の適合性"
challengeType: 28
dashedName: challenge-197
---

# --description--

ドナーの血液型とレシピエントの血液型が与えられたとき、ドナーがレシピエントに血液を提供できるかどうかを判定してください。

各血液型は以下で構成されます：

- 1文字のアルファベット：`"A"`、`"B"`、`"AB"`、または `"O"`
- そして Rh因子：`"+"` または `"-"`

血液型は有効なアルファベットのいずれかにRh因子が続く形になります。例えば、`"AB+"` と `"O-"` は有効な血液型です。

アルファベットのルール：

- `"O"` は他のアルファベット型に提供できます。
- `"A"` は `"A"` と `"AB"` に提供できます。
- `"B"` は `"B"` と `"AB"` に提供できます。
- `"AB"` は `"AB"` のみへ提供できます。

Rhのルール：

- マイナス（`"-"`）は `"-"` と `"+"` の両方に提供できます。
- プラス（`"+"`）は `"+"` のみへ提供できます。

ドナーがレシピエントに血液を提供できるには、アルファベットのルールとRhのルールの両方を満たす必要があります。

# --hints--

`canDonate("B+", "B+")` は `true` を返すべきです。

```js
assert.isTrue(canDonate("B+", "B+"));
```

`canDonate("O-", "AB-")` は `true` を返すべきです。

```js
assert.isTrue(canDonate("O-", "AB-"));
```

`canDonate("O+", "A-")` は `false` を返すべきです。

```js
assert.isFalse(canDonate("O+", "A-"));
```

`canDonate("A+", "AB+")` は `true` を返すべきです。

```js
assert.isTrue(canDonate("A+", "AB+"));
```

`canDonate("A-", "B-")` は `false` を返すべきです。

```js
assert.isFalse(canDonate("A-", "B-"));
```

`canDonate("B-", "AB+")` は `true` を返すべきです。

```js
assert.isTrue(canDonate("B-", "AB+"));
```

`canDonate("B-", "A+")` は `false` を返すべきです。

```js
assert.isFalse(canDonate("B-", "A+"));
```

`canDonate("O-", "O+")` は `true` を返すべきです。

```js
assert.isTrue(canDonate("O-", "O+"));
```

`canDonate("O+", "O-")` は `false` を返すべきです。

```js
assert.isFalse(canDonate("O+", "O-"));
```

`canDonate("AB+", "AB-")` は `false` を返すべきです。

```js
assert.isFalse(canDonate("AB+", "AB-"));
```

# --seed--

## --seed-contents--

```js
function canDonate(donor, recipient) {

  return donor;
}
```

# --solutions--

```js
function canDonate(donor, recipient) {
  const donorType = donor.slice(0, -1);
  const donorRh = donor.slice(-1);

  const recipientType = recipient.slice(0, -1);
  const recipientRh = recipient.slice(-1);

  const aboCompatibility = {
    "O": ["A", "B", "AB", "O"],
    "A": ["A", "AB"],
    "B": ["B", "AB"],
    "AB": ["AB"]
  };

  const aboMatch = aboCompatibility[donorType].includes(recipientType);
  const rhMatch =
    donorRh === "-" || (donorRh === "+" && recipientRh === "+");

  return aboMatch && rhMatch;
}
```
