---
id: 698a1a73ade5ac0e19180fa3
title: "チャレンジ 200: 文字と数字のカウント"
challengeType: 28
dashedName: challenge-200
---

# --description--

文字列が与えられたとき、その中に含まれる文字と数字の数を数えたメッセージを返してください。

- 文字は `A-Z` と `a-z` です。
- 数字は `0-9` です。
- その他の文字はすべて無視してください。

`"The string has X letters and Y numbers."` を返してください。ここで `"X"` は文字の数、`"Y"` は数字の数です。どちらかの数が1の場合は、その項目の単数形を使ってください。例：`"1 letter"` の代わりに `"1 letters"`、`"1 number"` の代わりに `"1 numbers"` を使います。

# --hints--

`countLettersAndNumbers("helloworld123")` は `"The string has 10 letters and 3 numbers."` を返すべきです。

```js
assert.equal(countLettersAndNumbers("helloworld123"), "The string has 10 letters and 3 numbers.");
```

`countLettersAndNumbers("Catch 22")` は `"The string has 5 letters and 2 numbers."` を返すべきです。

```js
assert.equal(countLettersAndNumbers("Catch 22"), "The string has 5 letters and 2 numbers.");
```

`countLettersAndNumbers("A1!")` は `"The string has 1 letter and 1 number."` を返すべきです。

```js
assert.equal(countLettersAndNumbers("A1!"), "The string has 1 letter and 1 number.");
```

`countLettersAndNumbers("12345")` は `"The string has 0 letters and 5 numbers."` を返すべきです。

```js
assert.equal(countLettersAndNumbers("12345"), "The string has 0 letters and 5 numbers.");
```

`countLettersAndNumbers("password")` は `"The string has 8 letters and 0 numbers."` を返すべきです。

```js
assert.equal(countLettersAndNumbers("password"), "The string has 8 letters and 0 numbers.");
```

# --seed--

## --seed-contents--

```js
function countLettersAndNumbers(str) {

  return str;
}
```

# --solutions--

```js
function countLettersAndNumbers(str) {
  let letterCount = 0;
  let numberCount = 0;

  for (let char of str) {
    if ((char >= 'A' && char <= 'Z') || (char >= 'a' && char <= 'z')) {
      letterCount++;
    } else if (char >= '0' && char <= '9') {
      numberCount++;
    }
  }

  const letterWord = letterCount === 1 ? "letter" : "letters";
  const numberWord = numberCount === 1 ? "number" : "numbers";

  return `The string has ${letterCount} ${letterWord} and ${numberCount} ${numberWord}.`;
}
```
