---
id: 698a1a863194f1f4e63f6460
title: "チャレンジ 209: 要素のサイズ"
challengeType: 28
dashedName: challenge-209
---

# --description--

ウィンドウサイズ、ビューポート幅の `"vw"` 単位での要素の幅、およびビューポート高さの `"vh"` 単位での要素の高さが与えられたとき、要素のサイズをピクセル単位で求めてください。

- 与えられるウィンドウサイズと返される要素サイズは、例えば `"width x height"`、`"1200 x 800"` の形式の文字列です。

- `"vw"` 単位はウィンドウ幅のパーセントです。例えば `"50vw"` はウィンドウ幅の 50% です。
- `"vh"` 単位はウィンドウ高さのパーセントです。例えば `"50vh"` はウィンドウ高さの 50% です。

# --hints--

`getElementSize("1200 x 800", "50vw", "50vh")` は `"600 x 400"` を返すべきです。

```js
assert.equal(getElementSize("1200 x 800", "50vw", "50vh"), "600 x 400");
```

`getElementSize("320 x 480", "25vw", "50vh")` は `"80 x 240"` を返すべきです。

```js
assert.equal(getElementSize("320 x 480", "25vw", "50vh"), "80 x 240");
```

`getElementSize("1000 x 500", "7vw", "3vh")` は `"70 x 15"` を返すべきです。

```js
assert.equal(getElementSize("1000 x 500", "7vw", "3vh"), "70 x 15");
```

`getElementSize("1920 x 1080", "95vw", "100vh")` は `"1824 x 1080"` を返すべきです。

```js
assert.equal(getElementSize("1920 x 1080", "95vw", "100vh"), "1824 x 1080");
```

`getElementSize("1200 x 800", "0vw", "0vh")` は `"0 x 0"` を返すべきです。

```js
assert.equal(getElementSize("1200 x 800", "0vw", "0vh"), "0 x 0");
```

`getElementSize("1440 x 900", "100vw", "114vh")` は `"1440 x 1026"` を返すべきです。

```js
assert.equal(getElementSize("1440 x 900", "100vw", "114vh"), "1440 x 1026");
```

# --seed--

## --seed-contents--

```js
function getElementSize(windowSize, elementVw, elementVh) {

  return windowSize;
}
```

# --solutions--

```js
function getElementSize(windowSize, elementVw, elementVh) {
  const [windowWidth, windowHeight] = windowSize
    .split(" x ")
    .map(Number);

  const vw = Number(elementVw.replace("vw", ""));
  const vh = Number(elementVh.replace("vh", ""));

  const widthPx = (vw / 100) * windowWidth;
  const heightPx = (vh / 100) * windowHeight;

  return `${widthPx} x ${heightPx}`;
}
```
