---
id: 698a1a73ade5ac0e19180fa6
title: "チャレンジ 203: フラットかどうか判定する"
challengeType: 28
dashedName: challenge-203
---

# --description--

配列がフラットかどうか判定してください。

- 配列の要素に配列が含まれていなければ、その配列はフラットです。

# --hints--

`isFlat([1, 2, 3, 4])` は `true` を返すべきです。

```js
assert.isTrue(isFlat([1, 2, 3, 4]));
```

`isFlat([1, [2, 3], 4])` は `false` を返すべきです。

```js
assert.isFalse(isFlat([1, [2, 3], 4]));
```

`isFlat([1, 0, false, true, "a", "b"])` は `true` を返すべきです。

```js
assert.isTrue(isFlat([1, 0, false, true, "a", "b"]));
```

`isFlat(["a", [0], "b", true])` は `false` を返すべきです。

```js
assert.isFalse(isFlat(["a", [0], "b", true]));
```

`isFlat([1, [2, [3, [4, [5]]]], 6])` は `false` を返すべきです。

```js
assert.isFalse(isFlat([1, [2, [3, [4, [5]]]], 6]));
```

# --seed--

## --seed-contents--

```js
function isFlat(arr) {

  return arr;
}
```

# --solutions--

```js
function isFlat(arr) {
  for (let el of arr) {
    if (Array.isArray(el)) return false;
  }
  return true;
}
```
