---
id: 698a1a73ade5ac0e19180fa2
title: "チャレンジ 199: 連続差分"
challengeType: 28
dashedName: challenge-199
---

# --description--

数値の配列が与えられたとき、各数値から次の数値に到達するために必要な値を含む新しい配列を返してください。

- 最後の数値には次の数値がないため、`0`を使用してください。

例えば、`[1, 2, 4, 7]`が与えられたら、`[1, 2, 3, 0]`を返します。

# --hints--

`findDifferences([1, 2, 4, 7])`は`[1, 2, 3, 0]`を返すべきです。

```js
assert.deepEqual(findDifferences([1, 2, 4, 7]), [1, 2, 3, 0]);
```

`findDifferences([10, 15, 19, 22, 24, 25])`は`[5, 4, 3, 2, 1, 0]`を返すべきです。

```js
assert.deepEqual(findDifferences([10, 15, 19, 22, 24, 25]), [5, 4, 3, 2, 1, 0]);
```

`findDifferences([25, 20, 16, 13, 11, 10])`は`[-5, -4, -3, -2, -1, 0]`を返すべきです。

```js
assert.deepEqual(findDifferences([25, 20, 16, 13, 11, 10]), [-5, -4, -3, -2, -1, 0]);
```

`findDifferences([0, 1, 2, 2, 3, 3, 4, 5])`は`[ 1, 1, 0, 1, 0, 1, 1, 0 ]`を返すべきです。

```js
assert.deepEqual(findDifferences([0, 1, 2, 2, 3, 3, 4, 5]), [ 1, 1, 0, 1, 0, 1, 1, 0 ]);
```

`findDifferences([1, 2, 5, 12, 34, -15, -1, 41, 113, -222, -99, -40, 10, -18, -6, -2, -1])`は`[1, 3, 7, 22, -49, 14, 42, 72, -335, 123, 59, 50, -28, 12, 4, 1, 0]`を返すべきです。

```js
assert.deepEqual(findDifferences([1, 2, 5, 12, 34, -15, -1, 41, 113, -222, -99, -40, 10, -18, -6, -2, -1]), [1, 3, 7, 22, -49, 14, 42, 72, -335, 123, 59, 50, -28, 12, 4, 1, 0]);
```

# --seed--

## --seed-contents--

```js
function findDifferences(arr) {

  return arr;
}
```

# --solutions--

```js
function findDifferences(arr) {
  const result = [];

  for (let i = 0; i < arr.length; i++) {
    result.push(arr[i + 1] - arr[i] || 0);
  }

  return result;
}
```
