---
id: 698a1a73ade5ac0e19180fa1
title: "チャレンジ 198: 営業日数の計算"
challengeType: 28
dashedName: challenge-198
---

# --description--

開始日と終了日が与えられたとき、その間の営業日数を返してください。

- 与えられる日付は `"YYYY-MM-DD"` の形式です。
- 平日（月曜日から金曜日）が営業日です。
- 週末（土曜日と日曜日）は営業日ではありません。
- 開始日と終了日もカウントに含めてください。

# --hints--

`countBusinessDays("2026-02-24", "2026-02-26")` は `3` を返すべきです。

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-26"), 3);
```

`countBusinessDays("2026-02-24", "2026-02-28")` は `4` を返すべきです。

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-28"), 4);
```

`countBusinessDays("2026-02-21", "2026-03-01")` は `5` を返すべきです。

```js
assert.equal(countBusinessDays("2026-02-21", "2026-03-01"), 5);
```

`countBusinessDays("2026-03-08", "2026-03-17")` は `7` を返すべきです。

```js
assert.equal(countBusinessDays("2026-03-08", "2026-03-17"), 7);
```

`countBusinessDays("2026-02-24", "2027-02-24")` は `262` を返すべきです。

```js
assert.equal(countBusinessDays("2026-02-24", "2027-02-24"), 262);
```

# --seed--

## --seed-contents--

```js
function countBusinessDays(start, end) {

  return start;
}
```

# --solutions--

```js
function countBusinessDays(start, end) {
  const startDate = new Date(start + "T00:00:00Z");
  const endDate = new Date(end + "T00:00:00Z");

  let count = 0;
  let current = new Date(startDate);

  while (current <= endDate) {
    const day = current.getUTCDay();
    if (day !== 0 && day !== 6) count++;
    current.setUTCDate(current.getUTCDate() + 1);
  }

  return count;
}
```
