---
id: 698a1a863194f1f4e63f645e
title: "チャレンジ 207: 最小のギャップ"
challengeType: 28
dashedName: challenge-207
---

# --description--

文字列が与えられたとき、最も文字数が少ない（最小のギャップ）同じ文字の間にある部分文字列を返してください。

- 常に少なくとも1組の一致する文字があります。
- 返す部分文字列には一致する文字自体は含めません。
- もし2つ以上のギャップが同じ長さの場合は、最初のものの文字を返してください。

例えば、`"ABCDAC"`が与えられた場合、`"DA"`を返します。理由は：

- 文字列内で繰り返されるのは`"A"`と`"C"`だけです。
- 2つの`"A"`文字の間の文字数は3で、`"C"`文字の間は2です。
- したがって、2つの`"C"`文字の間の文字列を返します。

# --hints--

`smallestGap("ABCDAC")`は`"DA"`を返すべきです。

```js
assert.equal(smallestGap("ABCDAC"), "DA");
```

`smallestGap("racecar")`は`"e"`を返すべきです。

```js
assert.equal(smallestGap("racecar"), "e");
```

`smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")`は`"#q6e&rkf(p"`を返すべきです。

```js
assert.equal(smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p");
```

`smallestGap("Hello World")`は`""`を返すべきです。

```js
assert.equal(smallestGap("Hello World"), "");
```

`smallestGap("The quick brown fox jumps over the lazy dog.")`は`"fox"`を返すべきです。

```js
assert.equal(smallestGap("The quick brown fox jumps over the lazy dog."), "fox");
```

# --seed--

## --seed-contents--

```js
function smallestGap(str) {

  return str;
}
```

# --solutions--

```js
function smallestGap(str) {
  let minGap = Infinity;
  let result = "";

  for (let i = 0; i < str.length; i++) {
    for (let j = i + 1; j < str.length; j++) {
      if (str[i] === str[j]) {
        const gap = j - i - 1;

        if (gap < minGap) {
          minGap = gap;
          result = str.slice(i + 1, j);
        }

        break;
      }
    }
  }

  return result;
}
```
