---
id: 698a1a73ade5ac0e19180fa8
title: "チャレンジ 205: 完全立方体の数を数える"
challengeType: 28
dashedName: challenge-205
---

# --description--

2つの整数が与えられたとき、その2つの数を含む範囲内に存在する完全立方体の数を求めてください。

- 小さい方の数が最初の引数であるとは限りません。
- 数が完全立方体であるとは、ある整数（`n`）が存在し、`n * n * n = number`となる場合を指します。例えば、27は`3 * 3 * 3 = 27`なので完全立方体です。

# --hints--

`countPerfectCubes(3, 30)`は`2`を返すべきです。

```js
assert.equal(countPerfectCubes(3, 30), 2);
```

`countPerfectCubes(1, 30)`は`3`を返すべきです。

```js
assert.equal(countPerfectCubes(1, 30), 3);
```

`countPerfectCubes(30, 0)`は`4`を返すべきです。

```js
assert.equal(countPerfectCubes(30, 0), 4);
```

`countPerfectCubes(-64, 64)`は`9`を返すべきです。

```js
assert.equal(countPerfectCubes(-64, 64), 9);
```

`countPerfectCubes(9214, -8127)`は`41`を返すべきです。

```js
assert.equal(countPerfectCubes(9214, -8127), 41);
```

# --seed--

## --seed-contents--

```js
function countPerfectCubes(a, b) {

  return a;
}
```

# --solutions--

```js
function countPerfectCubes(a, b) {
  const start = Math.min(a, b);
  const end = Math.max(a, b);

  let count = 0;
  let n = 0;

  while (n ** 3 <= end) {
    if (n ** 3 >= start) count++;
    n++;
  }

  n = -1;
  while (n ** 3 >= start) {
    if (n ** 3 <= end) count++;
    n--;
  }

  return count;
}
```
