---
id: 698a1a73ade5ac0e19180fa7
title: "チャレンジ 204: 文字の合計を求める"
challengeType: 28
dashedName: challenge-204
---

# --description--

文字列が与えられたら、その文字の合計値を返してください。

- 文字は大文字または小文字のA-Zです
- 文字の値は `"A"` = 1、`"B"` = 2、...、`"Z"` = 26 です
- 大文字と小文字の文字は同じ値を持ちます
- 文字以外のすべての文字は無視してください

# --hints--

`sumLetters("Hello")` は `52` を返すべきです。

```js
assert.equal(sumLetters("Hello"), 52);
```

`sumLetters("freeCodeCamp")` は `94` を返すべきです。

```js
assert.equal(sumLetters("freeCodeCamp"), 94);
```

`sumLetters("The quick brown fox jumps over the lazy dog.")` は `473` を返すべきです。

```js
assert.equal(sumLetters("The quick brown fox jumps over the lazy dog."), 473);
```

`sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci.")` は `1681` を返すべきです。

```js
assert.equal(sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci."), 1681);
```

`sumLetters("</404>")` は `0` を返すべきです。

```js
assert.equal(sumLetters("</404>"), 0);
```

# --seed--

## --seed-contents--

```js
function sumLetters(str) {

  return str;
}
```

# --solutions--

```js
function sumLetters(str) {
  let total = 0;
  for (let char of str) {
    let upper = char.toUpperCase();
    if (upper >= "A" && upper <= "Z") {
      total += upper.charCodeAt(0) - "A".charCodeAt(0) + 1;
    }
  }
  return total;
}
```
