---
id: 698a1a73ade5ac0e19180fa9
title: "チャレンジ 206: トランプの値"
challengeType: 28
dashedName: challenge-206
---

# --description--

トランプの配列が与えられたとき、各カードの数値を含む新しい配列を返してください。

カードの値:

- エース (`"A"`) は値が1です。
- 数字カード (`"2"` - `"10"`) はそれぞれの表面の数字の値、2から10です。
- フェイスカード：ジャック (`"J"`)、クイーン (`"Q"`)、キング (`"K"`) はそれぞれ10の値です。

スート（マーク）:

- 各カードにはスートがあります：スペード (`"S"`)、クラブ (`"C"`)、ダイヤ (`"D"`)、またはハート (`"H"`)。

カードの形式:

- 各カードは文字列で表されます：`"valueSuit"`。例えば、`"AS"` はスペードのエース、`"10H"` はハートの10、`"QC"` はクラブのクイーンです。

# --hints--

`cardValues(["3H", "4D", "5S"])` は `[3, 4, 5]` を返すべきです。

```js
assert.deepEqual(cardValues(["3H", "4D", "5S"]), [3, 4, 5]);
```

`cardValues(["AS", "10S", "10H", "6D", "7D"])` は `[1, 10, 10, 6, 7]` を返すべきです。

```js
assert.deepEqual(cardValues(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7]);
```

`cardValues(["8D", "QS", "2H", "JC", "9C"])` は `[8, 10, 2, 10, 9]` を返すべきです。

```js
assert.deepEqual(cardValues(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9]);
```

`cardValues(["AS", "KS"])` は `[1, 10]` を返すべきです。

```js
assert.deepEqual(cardValues(["AS", "KS"]), [1, 10]);
```

`cardValues(["10H", "JH", "QH", "KH", "AH"])` は `[10, 10, 10, 10, 1]` を返すべきです。

```js
assert.deepEqual(cardValues(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1]);
```

# --seed--

## --seed-contents--

```js
function cardValues(cards) {

  return cards;
}
```

# --solutions--

```js
function cardValues(cards) {
  const values = [];

  for (let card of cards) {
    let valueStr = card.slice(0, card.length - 1);

    let value;
    if (valueStr === "A") {
      value = 1;
    } else if (["J", "Q", "K"].includes(valueStr)) {
      value = 10;
    } else {
      value = parseInt(valueStr, 10);
    }

    values.push(value);
  }

  return values;
}
```
