---
id: 6994cff2290543b3aec9f510
title: "チャレンジ 211: 配列の合計"
challengeType: 28
dashedName: challenge-211
---

# --description--

数値の配列が与えられたら、すべての数値の合計を返してください。

# --hints--

`sumArray([1, 2, 3, 4, 5])` は `15` を返す必要があります。

```js
assert.equal(sumArray([1, 2, 3, 4, 5]), 15);
```

`sumArray([42])` は `42` を返す必要があります。

```js
assert.equal(sumArray([42]), 42);
```

`sumArray([5, -2, 7, -3])` は `7` を返す必要があります。

```js
assert.equal(sumArray([5, -2, 7, -3]), 7);
```

`sumArray([203, 145, -129, 6293, 523, -919, 845, 2434])` は `9395` を返す必要があります。

```js
assert.equal(sumArray([203, 145, -129, 6293, 523, -919, 845, 2434]), 9395);
```

`sumArray([0, 0])` は `0` を返す必要があります。

```js
assert.equal(sumArray([0, 0]), 0);
```

# --seed--

## --seed-contents--

```js
function sumArray(numbers) {

  return numbers;
}
```

# --solutions--

```js
function sumArray(numbers) {

  return numbers.reduce((a, c) => c + a, 0);
}
```
