---
id: 6994cff2290543b3aec9f50f
title: "チャレンジ 210: HSL バリデーター"
challengeType: 28
dashedName: challenge-210
---

# --description--

文字列が有効な CSS `hsl()` カラー値かどうかを判定してください。

- 有効な HSL 値は `"hsl(h, s%, l%)"` の形式である必要があります。具体的には：
  - `h`（色相）は 0 から 360（両端含む）の数値でなければなりません。
  - `s`（彩度）は 0% から 100% のパーセンテージでなければなりません。
  - `l`（明度）は 0% から 100% のパーセンテージでなければなりません。

- スペースは以下の場所にのみ許可されます：
  - 開き括弧の直後
  - カンマの前後いずれかまたは両方
  - 閉じ括弧の前後いずれかまたは両方

- オプションで、値の末尾にセミコロン（`";"`）を付けることができます。

例えば、`"hsl(240, 50%, 50%)"` は有効な HSL 値です。

# --hints--

`isValidHSL("hsl(240, 50%, 50%)")` は `true` を返すべきです。

```js
assert.isTrue(isValidHSL("hsl(240, 50%, 50%)"));
```

`isValidHSL("hsl( 200 , 50% , 75% )")` は `true` を返すべきです。

```js
assert.isTrue(isValidHSL("hsl( 200 , 50% , 75% )"));
```

`isValidHSL("hsl(99,60%,80%);")` は `true` を返すべきです。

```js
assert.isTrue(isValidHSL("hsl(99,60%,80%);"));
```

`isValidHSL("hsl(0, 0%, 0%) ;")` は `true` を返すべきです。

```js
assert.isTrue(isValidHSL("hsl(0, 0%, 0%) ;"));
```

`isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;")` は `true` を返すべきです。

```js
assert.isTrue(isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;"));
```

`isValidHSL("hsl(361, 50%, 80%)")` は `false` を返すべきです。

```js
assert.isFalse(isValidHSL("hsl(361, 50%, 80%)"));
```

`isValidHSL("hsl(300, 101%, 70%)")` は `false` を返すべきです。

```js
assert.isFalse(isValidHSL("hsl(300, 101%, 70%)"));
```

`isValidHSL("hsl(200, 55%, 75)")` は `false` を返すべきです。

```js
assert.isFalse(isValidHSL("hsl(200, 55%, 75)"));
```

`isValidHSL("hsl (80, 20%, 10%)")` は `false` を返すべきです。

```js
assert.isFalse(isValidHSL("hsl (80, 20%, 10%)"));
```

# --seed--

## --seed-contents--

```js
function isValidHSL(hsl) {

  return hsl;
}
```

# --solutions--

```js
function isValidHSL(hsl) {
  const regex = /^hsl\(\s*(\d{1,3})\s*,\s*(\d{1,3})%\s*,\s*(\d{1,3})%\s*\)\s*;?$/;

  const match = hsl.match(regex);
  if (!match) return false;

  const hue = Number(match[1]);
  const sat = Number(match[2]);
  const light = Number(match[3]);

  if (hue < 0 || hue > 360) return false;
  if (sat < 0 || sat > 100) return false;
  if (light < 0 || light > 100) return false;

  return true;
}
```
