---
id: 6994cff2290543b3aec9f511
title: "チャレンジ 212: 配列への挿入"
challengeType: 28
dashedName: challenge-212
---

# --description--

配列と、配列に挿入する値、そしてその値を挿入するインデックスが与えられたとき、指定されたインデックスに値が挿入された新しい配列を返してください。

# --hints--

`insertIntoArray([2, 4, 8, 10], 6, 2)` は `[2, 4, 6, 8, 10]` を返すべきです。

```js
assert.deepEqual(insertIntoArray([2, 4, 8, 10], 6, 2), [2, 4, 6, 8, 10]);
```

`insertIntoArray(["the", "quick", "fox"], "brown", 2)` は `["the", "quick", "brown", "fox"]` を返すべきです。

```js
assert.deepEqual(insertIntoArray(["the", "quick", "fox"], "brown", 2), ["the", "quick", "brown", "fox"]);
```

`insertIntoArray([], 0, 0)` は `[0]` を返すべきです。

```js
assert.deepEqual(insertIntoArray([], 0, 0), [0]);
```

`insertIntoArray([0, 1, 1, 2, 3, 8, 13], 5, 5)` は `[0, 1, 1, 2, 3, 5, 8, 13]` を返すべきです。

```js
assert.deepEqual(insertIntoArray([0, 1, 1, 2, 3, 8, 13], 5, 5), [0, 1, 1, 2, 3, 5, 8, 13]);
```

# --seed--

## --seed-contents--

```js
function insertIntoArray(arr, value, index) {

  return arr;
}
```

# --solutions--

```js
function insertIntoArray(arr, value, index) {
  return [
    ...arr.slice(0, index),
    value,
    ...arr.slice(index)
  ];
}
```
