---
id: 699c8e045ee7cb94ed2322db
title: "チャレンジ 220: 最大の数"
challengeType: 28
dashedName: challenge-220
---

# --description--

さまざまな句読点で区切られた数字の文字列が与えられたとき、最大の数を返してください。

- 与えられる文字列には数字と区切り文字のみが含まれます。
- 区切り文字はカンマ (`","`)、感嘆符 (`"!"`)、疑問符 (`"?"`)、コロン (`":"`)、セミコロン (`";"`) のいずれかです。

# --hints--

`largestNumber("1,2")` は `2` を返すべきです。

```js
assert.equal(largestNumber("1,2"), 2);
```

`largestNumber("4;15:60,26?52!0")` は `60` を返すべきです。

```js
assert.equal(largestNumber("4;15:60,26?52!0"), 60);
```

`largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011")` は `-402` を返すべきです。

```js
assert.equal(largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011"), -402);
```

`largestNumber("12;-50,99.9,49.1!-10.1?88?16")` は `99.9` を返すべきです。

```js
assert.equal(largestNumber("12;-50,99.9,49.1!-10.1?88?16"), 99.9);
```

# --seed--

## --seed-contents--

```js
function largestNumber(str) {

  return str;
}
```

# --solutions--

```js
function largestNumber(str) {
  const numbers = str.split(/[,!?:;]/).map(Number);
  return Math.max(...numbers);
}
```
