---
id: 699c8e045ee7cb94ed2322dd
title: "チャレンジ 222: 春分の影"
challengeType: 28
dashedName: challenge-222
---

# --description--

今日は春分の日で、太陽が赤道の真上にあり、正午には真上にあります。ある時刻が与えられたとき、4フィートの垂直な棒が作る影を求めてください。

- 時刻は`"HH:MM"`の24時間表記の文字列です（例えば、`"15:00"`は午後3時です）。
- 時刻は30分刻みでのみ与えられます。

ルール：

- 太陽は午前6時に`"east"`から昇り、午後6時に`"west"`に沈みます。
- 影は常に太陽の反対方向を指します。
- 影の長さ（フィート単位）は正午からの時間差の3乗です。
- 日の出前（午前6時前）、日没後（午後6時以降）、および正午には影はありません。

返す値：

- 影が存在する場合は`"(length)ft (direction)"`を返してください。例えば、`"8ft west"`です。
- それ以外の場合は`"No shadow"`を返してください。

例えば、`"10:00"`が与えられた場合、午前10時は正午から2時間離れているので2<sup>3</sup> = 8フィートで、太陽が東にあるため影は西を指すので`"8ft west"`を返します。

# --hints--

`getShadow("10:00")`は`"8ft west"`を返すべきです。

```js
assert.equal(getShadow("10:00"), "8ft west");
```

`getShadow("15:00")`は`"27ft east"`を返すべきです。

```js
assert.equal(getShadow("15:00"), "27ft east");
```

`getShadow("12:00")`は`"No shadow"`を返すべきです。

```js
assert.equal(getShadow("12:00"), "No shadow");
```

`getShadow("17:30")`は`"166.375ft east"`を返すべきです。

```js
assert.equal(getShadow("17:30"), "166.375ft east");
```

`getShadow("05:00")`は`"No shadow"`を返すべきです。

```js
assert.equal(getShadow("05:00"), "No shadow");
```

`getShadow("06:00")`は`"216ft west"`を返すべきです。

```js
assert.equal(getShadow("06:00"), "216ft west");
```

`getShadow("18:00")`は`"No shadow"`を返すべきです。

```js
assert.equal(getShadow("18:00"), "No shadow");
```

`getShadow("07:30")`は`"91.125ft west"`を返すべきです。

```js
assert.equal(getShadow("07:30"), "91.125ft west");
```

`getShadow("00:00")`は`"No shadow"`を返すべきです。

```js
assert.equal(getShadow("00:00"), "No shadow");
```

# --seed--

## --seed-contents--

```js
function getShadow(time) {

  return time;
}
```

# --solutions--

```js
function getShadow(time) {
  const [hourStr, minuteStr] = time.split(":");
  const hours = Number(hourStr);
  const minutes = Number(minuteStr);
  const timeInHours = hours + minutes / 60;

  if (timeInHours < 6 || timeInHours >= 18 || timeInHours === 12) {
    return "No shadow";
  }

  const hoursFromNoon = Math.abs(12 - timeInHours);
  const length = Math.pow(hoursFromNoon, 3);
  const direction = timeInHours < 12 ? "west" : "east";

  return `${length}ft ${direction}`;
}
```
