---
id: 699c8e045ee7cb94ed2322d5
title: "チャレンジ 214: ドミノ連鎖バリデーター"
challengeType: 28
dashedName: challenge-214
---

# --description--

2次元配列で表されたドミノのシーケンスが、有効な連鎖を形成しているかどうかを判定してください。

- 配列の各要素はドミノを表し、1から6までの数字2つの配列です（両端を含む）。
- 連鎖が有効であるためには、各ドミノの2番目の数字が次のドミノの1番目の数字と一致している必要があります。
- 最初のドミノの1番目の数字と最後のドミノの2番目の数字は、何とも一致している必要はありません。

# --hints--

`isValidDominoChain([[1, 3], [3, 6], [6, 5]])` は `true` を返すべきです。

```js
assert.isTrue(isValidDominoChain([[1, 3], [3, 6], [6, 5]]));
```

`isValidDominoChain([[6, 2], [3, 4], [4, 1]])` は `false` を返すべきです。

```js
assert.isFalse(isValidDominoChain([[6, 2], [3, 4], [4, 1]]));
```

`isValidDominoChain([[2, 5], [5, 6], [5, 1]])` は `false` を返すべきです。

```js
assert.isFalse(isValidDominoChain([[2, 5], [5, 6], [5, 1]]));
```

`isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]])` は `true` を返すべきです。

```js
assert.isTrue(isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]]));
```

`isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]])` は `false` を返すべきです。

```js
assert.isFalse(isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]]));
```

# --seed--

## --seed-contents--

```js
function isValidDominoChain(dominoes) {

  return dominoes;
}
```

# --solutions--

```js
function isValidDominoChain(dominoes) {
  for (let i = 0; i < dominoes.length - 1; i++) {
    if (dominoes[i][1] !== dominoes[i+1][0]) return false
  }

  return true;
}
```
