---
id: 699c8e045ee7cb94ed2322dc
title: "チャレンジ 221: 反転行列"
challengeType: 28
dashedName: challenge-221
---

# --description--

2つの異なる値で埋められた行列（配列の配列）が与えられたとき、一方の値のすべての出現箇所をもう一方の値と入れ替えた新しい行列を返してください。

例えば、次のような行列が与えられた場合：

```js
[
  ["a", "b"],
  ["a", "a"]
]
```

次のように返します：

```js
[
  ["b", "a"],
  ["b", "b"]
]
```

# --hints--

`invertMatrix([["a", "b"], ["a", "a"]])` は `[["b", "a"], ["b", "b"]]` を返すべきです。

```js
assert.deepEqual(invertMatrix([["a", "b"], ["a", "a"]]), [["b", "a"], ["b", "b"]]);
```

`invertMatrix([[1, 0, 1], [1, 1, 1], [0, 1, 0]])` は `[[0, 1, 0], [0, 0, 0], [1, 0, 1]]` を返すべきです。

```js
assert.deepEqual(invertMatrix([[1, 0, 1], [1, 1, 1], [0, 1, 0]]), [[0, 1, 0], [0, 0, 0], [1, 0, 1]]);
```

`invertMatrix([["apple", "banana", "banana", "apple"], ["banana", "apple", "apple", "banana"], ["banana", "banana", "banana", "apple"]])` は `[["banana", "apple", "apple", "banana"], ["apple", "banana", "banana", "apple"], ["apple", "apple", "apple", "banana"]]` を返すべきです。

```js
assert.deepEqual(invertMatrix([["apple", "banana", "banana", "apple"], ["banana", "apple", "apple", "banana"], ["banana", "banana", "banana", "apple"]]), [["banana", "apple", "apple", "banana"], ["apple", "banana", "banana", "apple"], ["apple", "apple", "apple", "banana"]]);
```

`invertMatrix([[6, 7, 7, 7, 6], [7, 6, 7, 6, 7], [7, 7, 6, 7, 7], [7, 6, 7, 6, 7], [6, 7, 7, 7, 6]])` は `[[7, 6, 6, 6, 7], [6, 7, 6, 7, 6], [6, 6, 7, 6, 6], [6, 7, 6, 7, 6], [7, 6, 6, 6, 7]]` を返すべきです。

```js
assert.deepEqual(invertMatrix([[6, 7, 7, 7, 6], [7, 6, 7, 6, 7], [7, 7, 6, 7, 7], [7, 6, 7, 6, 7], [6, 7, 7, 7, 6]]), [[7, 6, 6, 6, 7], [6, 7, 6, 7, 6], [6, 6, 7, 6, 6], [6, 7, 6, 7, 6], [7, 6, 6, 6, 7]]);
```

`invertMatrix([[1.2, 2.1, 2.1, 2.1], [2.1, 1.2, 2.1, 1.2], [1.2, 1.2, 2.1, 2.1]])` は `[[2.1, 1.2, 1.2, 1.2], [1.2, 2.1, 1.2, 2.1], [2.1, 2.1, 1.2, 1.2]]` を返すべきです。

```js
assert.deepEqual(invertMatrix([[1.2, 2.1, 2.1, 2.1], [2.1, 1.2, 2.1, 1.2], [1.2, 1.2, 2.1, 2.1]]), [[2.1, 1.2, 1.2, 1.2], [1.2, 2.1, 1.2, 2.1], [2.1, 2.1, 1.2, 1.2]]);
```

# --seed--

## --seed-contents--

```js
function invertMatrix(matrix) {

  return matrix;
}
```

# --solutions--

```js
function invertMatrix(matrix) {
  if (matrix.length === 0) return [];

  const values = [...new Set(matrix.flat())];
  const [val1, val2] = values;

  return matrix.map(row =>
    row.map(cell => (cell === val1 ? val2 : val1))
  );
}
```
