---
id: 699c8e045ee7cb94ed2322da
title: "チャレンジ 219: 記念日の節目"
challengeType: 28
dashedName: challenge-219
---

# --description--

結婚してからの年数を表す整数が与えられたとき、このチャートに従って直近の記念日の節目を返してください。

| 結婚年数 | 節目          |
| ------------- | ------------ |
| 1             | `"Paper"`    |
| 5             | `"Wood"`     |
| 10            | `"Tin"`      |
| 25            | `"Silver"`   |
| 40            | `"Ruby"`     |
| 50            | `"Gold"`     |
| 60            | `"Diamond"`  |
| 70            | `"Platinum"` |

- 最初の節目に達していない場合は `"Newlyweds"` を返してください。

# --hints--

`getMilestone(0)` は `"Newlyweds"` を返すべきです。

```js
assert.equal(getMilestone(0), "Newlyweds");
```

`getMilestone(1)` は `"Paper"` を返すべきです。

```js
assert.equal(getMilestone(1), "Paper");
```

`getMilestone(8)` は `"Wood"` を返すべきです。

```js
assert.equal(getMilestone(8), "Wood");
```

`getMilestone(10)` は `"Tin"` を返すべきです。

```js
assert.equal(getMilestone(10), "Tin");
```

`getMilestone(26)` は `"Silver"` を返すべきです。

```js
assert.equal(getMilestone(26), "Silver");
```

`getMilestone(45)` は `"Ruby"` を返すべきです。

```js
assert.equal(getMilestone(45), "Ruby");
```

`getMilestone(50)` は `"Gold"` を返すべきです。

```js
assert.equal(getMilestone(50), "Gold");
```

`getMilestone(64)` は `"Diamond"` を返すべきです。

```js
assert.equal(getMilestone(64), "Diamond");
```

`getMilestone(71)` は `"Platinum"` を返すべきです。

```js
assert.equal(getMilestone(71), "Platinum");
```

# --seed--

## --seed-contents--

```js
function getMilestone(years) {

  return years;
}
```

# --solutions--

```js
function getMilestone(years) {
  if (years < 1) return "Newlyweds";

  const milestones = [
    [1, "Paper"],
    [5, "Wood"],
    [10, "Tin"],
    [25, "Silver"],
    [40, "Ruby"],
    [50, "Gold"],
    [60, "Diamond"],
    [70, "Platinum"]
  ];

  for (let i = milestones.length - 1; i >= 0; i--) {
    if (years >= milestones[i][0]) {
      return milestones[i][1];
    }
  }
}
```
