---
id: 699c8e045ee7cb94ed2322d4
title: "チャレンジ 213: 単語の長さ変換器"
challengeType: 28
dashedName: challenge-213
---

# --description--

単語の文字列が与えられたら、各単語をその長さに置き換えた新しい文字列を返してください。

- 与えられた文字列の単語は、単一のスペースで区切られています
- 返す文字列のスペースはそのままにしてください。

例えば、`"hello world"` が与えられたら、`"5 5"` を返します。

# --hints--

`convertWords("hello world")` は `"5 5"` を返すべきです。

```js
assert.equal(convertWords("hello world"), "5 5");
```

`convertWords("Thanks and happy coding")` は `"6 3 5 6"` を返すべきです。

```js
assert.equal(convertWords("Thanks and happy coding"), "6 3 5 6");
```

`convertWords("The quick brown fox jumps over the lazy dog")` は `"3 5 5 3 5 4 3 4 3"` を返すべきです。

```js
assert.equal(convertWords("The quick brown fox jumps over the lazy dog"), "3 5 5 3 5 4 3 4 3");
```

`convertWords("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl")` は `"5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4"` を返すべきです。

```js
assert.equal(convertWords("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl"), "5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4");
```

# --seed--

## --seed-contents--

```js
function convertWords(str) {

  return str;
}
```

# --solutions--

```js
function convertWords(str) {
  return str
    .split(" ")
    .map(word => word.length)
    .join(" ");
}
```
