---
id: 699c8e045ee7cb94ed2322d9
title: "チャレンジ 218: きれいに割り切れるか"
challengeType: 28
dashedName: challenge-218
---

# --description--

2つの整数が与えられたとき、最初の整数が2番目の整数できれいに割り切れるかどうかを判定してください。

# --hints--

`isEvenlyDivisible(4, 2)` は `true` を返すべきです。

```js
assert.isTrue(isEvenlyDivisible(4, 2));
```

`isEvenlyDivisible(7, 3)` は `false` を返すべきです。

```js
assert.isFalse(isEvenlyDivisible(7, 3));
```

`isEvenlyDivisible(5, 10)` は `false` を返すべきです。

```js
assert.isFalse(isEvenlyDivisible(5, 10));
```

`isEvenlyDivisible(48, 6)` は `true` を返すべきです。

```js
assert.isTrue(isEvenlyDivisible(48, 6));
```

`isEvenlyDivisible(3186, 9)` は `true` を返すべきです。

```js
assert.isTrue(isEvenlyDivisible(3186, 9));
```

`isEvenlyDivisible(4192, 11)` は `false` を返すべきです。

```js
assert.isFalse(isEvenlyDivisible(4192, 11));
```

# --seed--

## --seed-contents--

```js
function isEvenlyDivisible(a, b) {

  return a;
}
```

# --solutions--

```js
function isEvenlyDivisible(a, b) {
  return a % b === 0;
}
```
