---
id: 699c8e045ee7cb94ed2322d6
title: "チャレンジ 215: 駐車料金計算機"
challengeType: 28
dashedName: challenge-215
---

# --description--

車を駐車した時間と引き取った時間を表す2つの文字列が与えられたとき、駐車料金を計算してください。

- 与えられる文字列は24時間制の`"HH:MM"`形式です。例えば`"14:00"`は午後2時を表します。
- 最初の文字列は車を駐車した時間、2つ目は引き取った時間です。
- 引き取り時間が駐車時間よりも早い場合は、駐車が深夜をまたいで翌日にまたがっていることを意味します。

料金ルール：

- 駐車1時間ごとに3ドルかかります。
- 端数の時間は切り上げて次の1時間として計算します。
- 駐車が深夜（午前0時）をまたぐ場合は、追加で10ドルの深夜料金がかかります。
- 最小料金は5ドルです（合計が5ドル未満の場合にのみ適用されます）。

合計料金を`"$cost"`、`"$5"`の形式で返してください。

# --hints--

`calculateParkingFee("09:00", "11:00")`は`"$6"`を返すべきです。

```js
assert.equal(calculateParkingFee("09:00", "11:00"), "$6");
```

`calculateParkingFee("10:00", "10:30")`は`"$5"`を返すべきです。

```js
assert.equal(calculateParkingFee("10:00", "10:30"), "$5");
```

`calculateParkingFee("08:10", "10:45")`は`"$9"`を返すべきです。

```js
assert.equal(calculateParkingFee("08:10", "10:45"), "$9");
```

`calculateParkingFee("14:40", "23:10")`は`"$27"`を返すべきです。

```js
assert.equal(calculateParkingFee("14:40", "23:10"), "$27");
```

`calculateParkingFee("18:15", "01:30")`は`"$34"`を返すべきです。

```js
assert.equal(calculateParkingFee("18:15", "01:30"), "$34");
```

`calculateParkingFee("11:11", "11:10")`は`"$82"`を返すべきです。

```js
assert.equal(calculateParkingFee("11:11", "11:10"), "$82");
```

# --seed--

## --seed-contents--

```js
function calculateParkingFee(parkTime, pickupTime) {

  return parkTime;
}
```

# --solutions--

```js
function calculateParkingFee(parkTime, pickupTime) {
  function toMinutes(time) {
    const [hours, minutes] = time.split(":").map(Number);
    return hours * 60 + minutes;
  }

  const entryMinutes = toMinutes(parkTime);
  const exitMinutes = toMinutes(pickupTime);

  let totalMinutes;
  let overnight = false;

  if (exitMinutes < entryMinutes) {
    overnight = true;
    totalMinutes = (24 * 60 - entryMinutes) + exitMinutes;
  } else {
    totalMinutes = exitMinutes - entryMinutes;
  }

  const hours = Math.ceil(totalMinutes / 60);
  let cost = hours * 3;

  if (overnight) {
    cost += 10;
  }

  if (cost < 5) {
    cost = 5;
  }

  return `$${cost}`;
}
```
