---
id: 699c8e045ee7cb94ed2322d8
title: "チャレンジ217: 捕獲されたチェスの駒"
challengeType: 28
dashedName: challenge-217
---

# --description--

チェスの盤上にまだある駒を表す文字列の配列が与えられたとき、相手に捕獲された駒の合計価値を計算してください。

チェスでは、最初に16個の駒を持っています：

| 駒     | 略称           | 数量     | 価値  |
| ------ | -------------- | -------- | ----- |
| ポーン | `"P"`          | 8        | 1     |
| ルーク | `"R"`          | 2        | 5     |
| ナイト | `"N"`          | 2        | 3     |
| ビショップ | `"B"`          | 2        | 3     |
| クイーン | `"Q"`          | 1        | 9     |
| キング | `"K"`          | 1        | 0     |

- 与えられた配列には上記の略称のみが含まれます。
- 16個の駒のうち、配列に含まれていないものはすべて捕獲されています。
- 捕獲されたすべての駒の合計価値を返してください。ただし…
- キングが捕獲されている場合は、`"Checkmate"`を返してください。

# --hints--

`getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"])`は`8`を返すべきです。

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"]), 8);
```

`getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"])`は`26`を返すべきです。

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"]), 26);
```

`getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"])`は`16`を返すべきです。

```js
assert.equal(getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"]), 16);
```

`getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"])`は`4`を返すべきです。

```js
assert.equal(getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"]), 4);
```

`getCapturedValue(["P", "K"])`は`38`を返すべきです。

```js
assert.equal(getCapturedValue(["P", "K"]), 38);
```

`getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"])`は`0`を返すべきです。

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"]), 0);
```

`getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"])`は`"Checkmate"`を返すべきです。

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"]), "Checkmate");
```

# --seed--

## --seed-contents--

```js
function getCapturedValue(pieces) {

  return pieces;
}
```

# --solutions--

```js
function getCapturedValue(pieces) {
  if (!pieces.includes("K")) return "Checkmate";

  const values = { P:1, R:5, N:3, B:3, Q:9, K:0 };
  let remaining = "PPPPPPPPRRNNBBQK";

  for (const p of pieces) {
    remaining = remaining.replace(p, "");
  }

  return [...remaining].reduce((sum, p) => sum + values[p], 0);
}
```
