---
id: 69a890af247de743333bd4ce
title: "チャレンジ225: 連続した繰り返しなし"
challengeType: 28
dashedName: challenge-225
---

# --description--

文字列が与えられたとき、連続した繰り返し文字がないかどうかを判定してください。

- 文字列に同じ文字が2回以上連続していなければ、繰り返しがないとみなします。

# --hints--

`hasNoRepeats("hi world")` は `true` を返すべきです。

```js
assert.isTrue(hasNoRepeats("hi world"));
```

`hasNoRepeats("hello world")` は `false` を返すべきです。

```js
assert.isFalse(hasNoRepeats("hello world"));
```

`hasNoRepeats("abcdefghijklmnopqrstuvwxyz")` は `true` を返すべきです。

```js
assert.isTrue(hasNoRepeats("abcdefghijklmnopqrstuvwxyz"));
```

`hasNoRepeats("freeCodeCamp")` は `false` を返すべきです。

```js
assert.isFalse(hasNoRepeats("freeCodeCamp"));
```

`hasNoRepeats("The quick brown fox jumped over the lazy dog.")` は `true` を返すべきです。

```js
assert.isTrue(hasNoRepeats("The quick brown fox jumped over the lazy dog."));
```

`hasNoRepeats("Mississippi")` は `false` を返すべきです。

```js
assert.isFalse(hasNoRepeats("Mississippi"));
```

# --seed--

## --seed-contents--

```js
function hasNoRepeats(str) {

  return str;
}
```

# --solutions--

```js
function hasNoRepeats(str) {
  for (let i = 1; i < str.length; i++) {
    if (str[i] === str[i - 1]) {
      return false;
    }
  }
  return true;
}
```
