---
id: 69a890af247de743333bd4d1
title: "チャレンジ 228: 映画の夜"
challengeType: 28
dashedName: challenge-228
---

# --description--

曜日を表す文字列、上映時間を表す別の文字列、そしてチケットの枚数を表す整数が与えられたとき、その上映の映画チケットの合計金額を返してください。

与えられる曜日は次のいずれかです：

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

上映時間は `"H:MMam"` または `"H:MMpm"` の形式で与えられます。例えば `"10:00am"` や `"10:00pm"` のようにです。

次のルールに従い、合計金額を `"$D.CC"` の形式で返してください：

- 週末（金曜日から日曜日）：チケット1枚あたり12.00ドル。
- 平日（月曜日から木曜日）：チケット1枚あたり10.00ドル。
- マチネ（午後5時前）：チケット1枚あたり2.00ドル割引（火曜日を除く）。
- 火曜日：すべてのチケットが1枚5.00ドル。

# --hints--

`getMovieNightCost("Saturday", "10:00pm", 1)` は `"$12.00"` を返すべきです。

```js
assert.equal(getMovieNightCost("Saturday", "10:00pm", 1), "$12.00");
```

`getMovieNightCost("Sunday", "10:00am", 1)` は `"$10.00"` を返すべきです。

```js
assert.equal(getMovieNightCost("Sunday", "10:00am", 1), "$10.00");
```

`getMovieNightCost("Tuesday", "7:20pm", 2)` は `"$10.00"` を返すべきです。

```js
assert.equal(getMovieNightCost("Tuesday", "7:20pm", 2), "$10.00");
```

`getMovieNightCost("Wednesday", "5:40pm", 3)` は `"$30.00"` を返すべきです。

```js
assert.equal(getMovieNightCost("Wednesday", "5:40pm", 3), "$30.00");
```

`getMovieNightCost("Monday", "11:50am", 4)` は `"$32.00"` を返すべきです。

```js
assert.equal(getMovieNightCost("Monday", "11:50am", 4), "$32.00");
```

`getMovieNightCost("Friday", "4:30pm", 5)` は `"$50.00"` を返すべきです。

```js
assert.equal(getMovieNightCost("Friday", "4:30pm", 5), "$50.00");
```

`getMovieNightCost("Tuesday", "11:30am", 1)` は `"$5.00"` を返すべきです。

```js
assert.equal(getMovieNightCost("Tuesday", "11:30am", 1), "$5.00");
```

# --seed--

## --seed-contents--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {

  return day;
}
```

# --solutions--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {
  let price;

  if (day === "Tuesday") {
    price = 5;
  } else {
    if (["Friday", "Saturday", "Sunday"].includes(day)) {
      price = 12;
    } else {
      price = 10;
    }

    const hour = parseInt(showtime.split(":")[0]);
    const period = showtime.slice(-2);

    let hour24 = hour;

    if (period === "pm" && hour !== 12) hour24 += 12;
    if (period === "am" && hour === 12) hour24 = 0;

    if (hour24 < 17) {
      price -= 2;
    }
  }

  const total = price * numberOfTickets;

  return `$${total.toFixed(2)}`;
}
```
