---
id: 69a890af247de743333bd4d2
title: "チャレンジ 229: 文字列の切り詰め 2"
challengeType: 28
dashedName: challenge-229
---

# --description--

文字列が与えられたとき、文字の合計幅が50単位を超えないように切り詰めた新しい文字列を返してください。

各文字には特定の幅があります：

| 文字 | 幅 |
| - | - |
| `"ilI"` | 1 |
| `"fjrt"` | 2 |
| `"abcdeghkmnopqrstuvwxyzJL"` | 3 |
| `"ABCDEFGHKMNOPQRSTUVWXYZ"` | 4 |

上のテーブルにはすべての大文字と小文字が含まれています。さらに：

- スペース（`" "`）は幅が2です
- ピリオド（`"."`）は幅が1です

- 与えられた文字列の幅が50単位以下の場合は、そのまま文字列を返してください。そうでなければ
- 文字列を切り詰め、末尾に3つのピリオド（`"..."`）を追加して、3つのピリオドを含めた合計幅が50単位にできるだけ近く、超えないようにしてください。

# --hints--

`truncateText("The quick brown fox")` は `"The quick brown f..."` を返すべきです。

```js
assert.equal(truncateText("The quick brown fox"), "The quick brown f...");
```

`truncateText("The silky smooth sloth")` は `"The silky smooth s..."` を返すべきです。

```js
assert.equal(truncateText("The silky smooth sloth"), "The silky smooth s...");
```

`truncateText("THE LOUD BRIGHT BIRD")` は `"THE LOUD BRIG..."` を返すべきです。

```js
assert.equal(truncateText("THE LOUD BRIGHT BIRD"), "THE LOUD BRIG...");
```

`truncateText("The fast striped zebra")` は `"The fast striped z..."` を返すべきです。

```js
assert.equal(truncateText("The fast striped zebra"), "The fast striped z...");
```

`truncateText("The big black bear")` は `"The big black bear"` を返すべきです。

```js
assert.equal(truncateText("The big black bear"), "The big black bear");
```

# --seed--

## --seed-contents--

```js
function truncateText(str) {

  return str;
}
```

# --solutions--

```js
function truncateText(str) {
  const MAX_WIDTH = 50;
  const ELLIPSIS = "...";
  const ELLIPSIS_WIDTH = 3;
  const TRUNCATE_LIMIT = MAX_WIDTH - ELLIPSIS_WIDTH;

  const charGroups = {
    'ilI.': 1,
    'fjrt ': 2,
    'abcdeghkmnopqrstuvwxyzJL': 3,
    'ABCDEFGHKMNOPQRSTUVWXYZ': 4
  };

  function getCharWidth(char) {
    for (const key in charGroups) {
      if (key.includes(char)) return charGroups[key];
    }
    return 3;
  }

  function stringWidth(str) {
    let totalWidth = 0;
    for (const char of str) {
      totalWidth += getCharWidth(char);
    }
    return totalWidth;
  }

  if (stringWidth(str) <= MAX_WIDTH) return str;

  let result = "";
  let totalWidth = 0;

  for (const char of str) {
    const charWidth = getCharWidth(char);

    if (totalWidth + charWidth > TRUNCATE_LIMIT) break;

    result += char;
    totalWidth += charWidth;
  }

  return result + ELLIPSIS;
}
```
