---
id: 69a890af247de743333bd4d0
title: "チャレンジ 227: クールダウン時間"
challengeType: 28
dashedName: challenge-227
---

# --description--

2つのタイムスタンプが与えられます。1つ目はユーザーが試験を終えた時刻、2つ目は現在の時刻を表します。ユーザーが再度試験を受けられるかどうかを判定してください。

- 両方のタイムスタンプは `"YYYY-MM-DDTHH:MM:SS"` の形式で与えられます。例えば `"2026-03-25T14:00:00"` のようになります。時間は24時間制です。
- ユーザーは試験を再受験する前に少なくとも48時間待つ必要があります。

# --hints--

`canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` は `true` を返すべきです。

```js
assert.isTrue(canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00"));
```

`canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` は `false` を返すべきです。

```js
assert.isFalse(canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00"));
```

`canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` は `true` を返すべきです。

```js
assert.isTrue(canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00"));
```

`canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59")` は `false` を返すべきです。

```js
assert.isFalse(canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59"));
```

# --seed--

## --seed-contents--

```js
function canRetake(finishTime, currentTime) {

  return finishTime;
}
```

# --solutions--

```js
function canRetake(finishTime, currentTime) {
  function parse(ts) {
    const parts = ts.split(/[-T:]/).map(Number);
    const [y, m, d, h, min, s] = parts;
    return Date.UTC(y, m - 1, d, h, min, s);
  }

  const finishMs = parse(finishTime);
  const currentMs = parse(currentTime);

  const cooldown = 48 * 60 * 60 * 1000;

  return currentMs - finishMs >= cooldown;
}
```
