---
id: 69a890af247de743333bd4cc
title: "チャレンジ 223: QRデコーダー"
challengeType: 28
dashedName: challenge-223
---

# --description--

6x6の行列（配列の配列）で表されたQRコードが与えられたとき、そのコード内の2進データの文字列を返してください。

- QRコードは90度単位で任意の回転で与えられる場合があります。
- 正しい向きのコードは、左下、左上、右上の角に2x2の`1`（方向マーカー）のグループがあります。
- 3つの2x2の方向マーカーは2進データの一部ではありません。
- QRコードが正しい向きの場合、2進データは左から右へ、上から下へ（本のように）読み取られます。
- コードには必ず1つだけ有効な向きがあります。

例えば、次のようなコードが与えられた場合：

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "110000",
  "110001"
]
```

または異なる向きで同じコードが与えられた場合：

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "000011",
  "100011"
]
```

3つの2x2方向マーカーを除いたすべての2進データ、すなわち`"000000000000000000000001"`を返してください。

# --hints--

`decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"])`は`"000000000000000000000001"`を返すべきです。

```js
assert.equal(decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"]), "000000000000000000000001");
```

`decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"])`は`"000000000000000000000001"`を返すべきです。

```js
assert.equal(decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"]), "000000000000000000000001");
```

`decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"])`は`"001101000011000000110100"`を返すべきです。

```js
assert.equal(decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"]), "001101000011000000110100");
```

`decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"])`は`"010001000100010101010110"`を返すべきです。

```js
assert.equal(decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"]), "010001000100010101010110");
```

`decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"])`は`"010000100100100101001110"`を返すべきです。

```js
assert.equal(decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"]), "010000100100100101001110");
```

# --seed--

## --seed-contents--

```js
function decodeQr(qrCode) {

  return qrCode;
}
```

# --solutions--

```js
function isOriented(code) {
  const markerIndices = [
    [0, 0],
    [0, 1],
    [0, 4],
    [0, 5],
    [1, 0],
    [1, 1],
    [1, 4],
    [1, 5],
    [4, 0],
    [4, 1],
    [5, 0],
    [5, 1]
  ];

  for(const [row, col] of markerIndices) {
    if (code[row][col] !== '1') return false;
  }

  return true;
}

function rotate(code) {
  const newCode = ['', '', '', '', '', ''];
  for (let i = 0; i < 6; i++) {
    for (let j = 5; j >= 0; j--) {
      newCode[i] += code[j][i];
    }
  }
  return newCode;
}

function getData(code) {
  let data = '';
  data += code[0].substring(2, 4)
  data += code[1].substring(2, 4)
  data += code[2]
  data += code[3]
  data += code[4].substring(2)
  data += code[5].substring(2)
  return data;
}

function decodeQr(qrCode) {
  let code = [...qrCode];

  while(!isOriented(code)) {
    code = rotate(code);
  }

  return getData(code);
}
```
