---
id: 69a890af247de743333bd4cd
title: "チャレンジ 224: コーヒーロースト検出器"
challengeType: 28
dashedName: challenge-224
---

# --description--

コーヒー1杯を作るために使われた豆を表す文字列が与えられたとき、そのコーヒーのローストを判定してください。

- 与えられる文字列には、各豆の種類を表す以下の文字が含まれます：

  - アポストロフィー (`'`) はライトロースト豆で、1ポイントの価値があります。
  - ダッシュ (`-`) はミディアムロースト豆で、2ポイントの価値があります。
  - ピリオド (`.`) はダークロースト豆で、3ポイントの価値があります。

- ローストレベルは、すべての豆の平均値によって決まります。

返す値：

- 平均が1.75未満の場合は `"Light"` を返してください。
- 平均が1.75から2.5の場合は `"Medium"` を返してください。
- 平均が2.5より大きい場合は `"Dark"` を返してください。

# --hints--

`detectRoast("''-''''''-'-''--''''")` は `"Light"` を返すべきです。

```js
assert.equal(detectRoast("''-''''''-'-''--''''"), "Light");
```

`detectRoast(".'-''-''..'''.-.-''-")` は `"Medium"` を返すべきです。

```js
assert.equal(detectRoast(".'-''-''..'''.-.-''-"), "Medium");
```

`detectRoast("--.''--'-''.--..-.--")` は `"Medium"` を返すべきです。

```js
assert.equal(detectRoast("--.''--'-''.--..-.--"), "Medium");
```

`detectRoast("-...'-......-..-...-")` は `"Dark"` を返すべきです。

```js
assert.equal(detectRoast("-...'-......-..-...-"), "Dark");
```

`detectRoast(".--.-..-......----.'")` は `"Medium"` を返すべきです。

```js
assert.equal(detectRoast(".--.-..-......----.'"), "Medium");
```

`detectRoast("..-..-..-..-....-.-.")` は `"Dark"` を返すべきです。

```js
assert.equal(detectRoast("..-..-..-..-....-.-."), "Dark");
```

`detectRoast("-'-''''''..-'.''-'.'")` は `"Light"` を返すべきです。

```js
assert.equal(detectRoast("-'-''''''..-'.''-'.'"), "Light");
```

# --seed--

## --seed-contents--

```js
function detectRoast(beans) {

  return beans;
}
```

# --solutions--

```js
function detectRoast(beans) {
  let total = 0;

  for (const bean of beans) {
    if (bean === "'") total += 1;
    else if (bean === "-") total += 2;
    else if (bean === ".") total += 3;
  }

  const avg = total / beans.length;

  if (avg < 1.75) return "Light";
  if (avg <= 2.5) return "Medium";
  return "Dark";
}
```
