---
id: 69b1028d6e265413d0198a2f
title: "チャレンジ 236: ブラウザ履歴"
challengeType: 28
dashedName: challenge-236
---

# --description--

ブラウザのコマンドの配列が与えられたとき、履歴をURLの配列として、現在のページのインデックスとともに2つの値を持つ配列を返してください。

有効なコマンドは以下の通りです：

- `"URL"` - URLはウェブアドレスです（例えば`"freecodecamp.org"`）。指定されたURLに移動し、履歴の次の位置に追加し、先の履歴は破棄します。
- `"Back"` - 履歴の前のページに移動します。前のページがなければ現在のページに留まります。
- `"Forward"` - 履歴の次のページに移動します。次のページがなければ現在のページに留まります。

例えば、`["freecodecamp.org", "freecodecamp.org/learn", "Back"]`が与えられたら、`[["freecodecamp.org", "freecodecamp.org/learn"], 0]`を返してください。

# --hints--

`getBrowserHistory(["freecodecamp.org", "freecodecamp.org/learn", "Back"])`は`[["freecodecamp.org", "freecodecamp.org/learn"], 0]`を返すはずです。

```js
assert.deepEqual(getBrowserHistory(["freecodecamp.org", "freecodecamp.org/learn", "Back"]), [["freecodecamp.org", "freecodecamp.org/learn"], 0]);
```

`getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog"])`は`[["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]`を返すはずです。

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog"]), [["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]);
```

`getBrowserHistory(["example.com", "example.com/about", "Back", "example.com/contact",  "example.com/blog", "Back", "Back", "Forward"])`は`[["example.com", "example.com/contact", "example.com/blog"], 1]`を返すはずです。

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Back", "example.com/contact", "example.com/blog", "Back", "Back", "Forward"]), [["example.com", "example.com/contact", "example.com/blog"], 1]);
```

`getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"])`は`[["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]`を返すはずです。

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"]), [["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]);
```

`getBrowserHistory(["example.com", "example.com/about", "Back", "Back"])`は`[["example.com", "example.com/about"], 0]`を返すはずです。

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Back", "Back"]), [["example.com", "example.com/about"], 0]);
```

`getBrowserHistory(["example.com", "example.com/about", "Forward"])`は`[["example.com", "example.com/about"], 1]`を返すはずです。

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Forward"]), [["example.com", "example.com/about"], 1]);
```

# --seed--

## --seed-contents--

```js
function getBrowserHistory(commands) {

  return commands;
}
```

# --solutions--

```js
function getBrowserHistory(commands) {
  const stack = [];
  let index = -1;

  for (const command of commands) {
    if (command === "Back") {
      if (index > 0) index--;
    } else if (command === "Forward") {
      if (index < stack.length - 1) index++;
    } else {
      stack.splice(index + 1);
      stack.push(command);
      index++;
    }
  }

  return [stack, index];
}
```
