---
id: 69b1028d6e265413d0198a2c
title: "チャレンジ 233: 起床アラーム"
challengeType: 28
dashedName: challenge-233
---

# --description--

アラームを設定した時間を表す文字列と、実際に起きた時間を表す文字列が与えられたとき、早く起きたのか、時間通りに起きたのか、それとも遅く起きたのかを判定してください。

- 両方の時間は`"HH:MM"` 24時間形式で与えられます。

返す値:

- アラーム時間より前に起きた場合は`"early"`を返してください。
- アラーム時間ちょうど、またはアラーム時間から10分以内のスヌーズ時間内に起きた場合は`"on time"`を返してください。
- アラーム時間から10分以上遅れて起きた場合は`"late"`を返してください。

両方の時間は同じ日のものです。

# --hints--

`alarmCheck("07:00", "06:45")`は`"early"`を返すべきです。

```js
assert.equal(alarmCheck("07:00", "06:45"), "early");
```

`alarmCheck("06:30", "06:30")`は`"on time"`を返すべきです。

```js
assert.equal(alarmCheck("06:30", "06:30"), "on time");
```

`alarmCheck("08:10", "08:15")`は`"on time"`を返すべきです。

```js
assert.equal(alarmCheck("08:10", "08:15"), "on time");
```

`alarmCheck("09:30", "09:45")`は`"late"`を返すべきです。

```js
assert.equal(alarmCheck("09:30", "09:45"), "late");
```

`alarmCheck("08:15", "08:25")`は`"on time"`を返すべきです。

```js
assert.equal(alarmCheck("08:15", "08:25"), "on time");
```

`alarmCheck("05:45", "05:56")`は`"late"`を返すべきです。

```js
assert.equal(alarmCheck("05:45", "05:56"), "late");
```

`alarmCheck("04:30", "04:00")`は`"early"`を返すべきです。

```js
assert.equal(alarmCheck("04:30", "04:00"), "early");
```

# --seed--

## --seed-contents--

```js
function alarmCheck(alarmTime, wakeTime) {

  return alarmTime;
}
```

# --solutions--

```js
function alarmCheck(alarmTime, wakeTime) {
  const [alarmH, alarmM] = alarmTime.split(":").map(Number);
  const [wakeH, wakeM] = wakeTime.split(":").map(Number);

  const alarmMinutes = alarmH * 60 + alarmM;
  const wakeMinutes = wakeH * 60 + wakeM;
  const diff = wakeMinutes - alarmMinutes;

  if (diff < 0) return "early";
  if (diff <= 10) return "on time";
  return "late";
}
```
