---
id: 69b1028d6e265413d0198a29
title: "チャレンジ 230: パスカルの三角形の行"
challengeType: 28
dashedName: challenge-230
---

# --description--

整数 `n` が与えられたとき、パスカルの三角形の `n` 行目を配列として返してください。

パスカルの三角形では、各行は1で始まり1で終わり、内部の値はその直上の2つの値の合計です。

以下は三角形の最初の5行です：

```js
    1
   1 1
  1 2 1
 1 3 3 1
1 4 6 4 1
```

# --hints--

`pascalRow(5)` は `[1, 4, 6, 4, 1]` を返すべきです。

```js
assert.deepEqual(pascalRow(5), [1, 4, 6, 4, 1]);
```

`pascalRow(3)` は `[1, 2, 1]` を返すべきです。

```js
assert.deepEqual(pascalRow(3), [1, 2, 1]);
```

`pascalRow(1)` は `[1]` を返すべきです。

```js
assert.deepEqual(pascalRow(1), [1]);
```

`pascalRow(10)` は `[1, 9, 36, 84, 126, 126, 84, 36, 9, 1]` を返すべきです。

```js
assert.deepEqual(pascalRow(10), [1, 9, 36, 84, 126, 126, 84, 36, 9, 1]);
```

`pascalRow(15)` は `[1, 14, 91, 364, 1001, 2002, 3003, 3432, 3003, 2002, 1001, 364, 91, 14, 1]` を返すべきです。

```js
assert.deepEqual(pascalRow(15), [1, 14, 91, 364, 1001, 2002, 3003, 3432, 3003, 2002, 1001, 364, 91, 14, 1]);
```

# --seed--

## --seed-contents--

```js
function pascalRow(n) {

  return n;
}
```

# --solutions--

```js
function pascalRow(n) {
  if (n === 1) return [1];

  let row = [1];

  for (let i = 2; i <= n; i++) {
    const next = [1];

    for (let j = 1; j < row.length; j++) {
      next.push(row[j - 1] + row[j]);
    }

    next.push(1);
    row = next;
  }

  return row;
}
```
