---
id: 69b1028d6e265413d0198a2a
title: "チャレンジ 231: ISBN-10 バリデーター"
challengeType: 28
dashedName: challenge-231
---

# --description--

文字列が有効なISBN-10かどうか判定してください。

ISBN-10はハイフン（`"-"`）とその他10文字で構成されます。ハイフンを取り除くと（`"-"`）：

- 最初の9文字は数字でなければならず、
- 最後の文字は数字か、数値の10を表す文字`"X"`である場合があります。

検証方法は：

- 各数字（または値）にその位置を乗算します（最初の数字に1を、2番目に2を、というように）。
- すべての結果を合計します。
- 合計が11で割り切れれば有効です。

# --hints--

`isValidIsbn10("0-306-40615-2")`は`true`を返すべきです。

```js
assert.isTrue(isValidIsbn10("0-306-40615-2"));
```

`isValidIsbn10("0-306-40615-1")`は`false`を返すべきです。

```js
assert.isFalse(isValidIsbn10("0-306-40615-1"));
```

`isValidIsbn10("0-8044-2957-X")`は`true`を返すべきです。

```js
assert.isTrue(isValidIsbn10("0-8044-2957-X"));
```

`isValidIsbn10("X-306-40615-2")`は`false`を返すべきです。

```js
assert.isFalse(isValidIsbn10("X-306-40615-2"));
```

`isValidIsbn10("0-6822-2589-4")`は`true`を返すべきです。

```js
assert.isTrue(isValidIsbn10("0-6822-2589-4"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn10(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn10(str) {
  const isbn = str.replace(/-/g, "");

  if (isbn.length !== 10) return false;

  let sum = 0;

  for (let i = 0; i < 9; i++) {
    const digit = Number(isbn[i]);
    if (Number.isNaN(digit)) return false;

    sum += digit * (i + 1);
  }

  const last = isbn[9];

  if (last === "X") {
    sum += 10 * 10;
  } else {
    const digit = Number(last);
    if (Number.isNaN(digit)) return false;

    sum += digit * 10;
  }

  return sum % 11 === 0;
}
```
