---
id: 69b1028d6e265413d0198a2d
title: "チャレンジ 234: いたずらな数字"
challengeType: 28
dashedName: challenge-234
---

# --description--

すべての数字があるパターンに従っている配列が与えられたとき、そのパターンに従っていない1つの数字を修正した新しい配列を返してください。

パターンは次のいずれかです：

- 数字が一定の量だけ増加している（加算）。
- 数字が一定の量だけ減少している（減算）。

例えば、`[2, 4, 7, 8, 10]` が与えられたら `[2, 4, 6, 8, 10]` を返してください。

# --hints--

`fixPrankNumber([2, 4, 7, 8, 10])` は `[2, 4, 6, 8, 10]` を返すべきです。

```js
assert.deepEqual(fixPrankNumber([2, 4, 7, 8, 10]), [2, 4, 6, 8, 10]);
```

`fixPrankNumber([10, 10, 8, 7, 6])` は `[10, 9, 8, 7, 6]` を返すべきです。

```js
assert.deepEqual(fixPrankNumber([10, 10, 8, 7, 6]), [10, 9, 8, 7, 6]);
```

`fixPrankNumber([12, 24, 36, 48, 61, 72, 84, 96])` は `[12, 24, 36, 48, 60, 72, 84, 96]` を返すべきです。

```js
assert.deepEqual(fixPrankNumber([12, 24, 36, 48, 61, 72, 84, 96]), [12, 24, 36, 48, 60, 72, 84, 96]);
```

`fixPrankNumber([4, 1, -2, -5, -8, -5])` は `[4, 1, -2, -5, -8, -11]` を返すべきです。

```js
assert.deepEqual(fixPrankNumber([4, 1, -2, -5, -8, -5]), [4, 1, -2, -5, -8, -11]);
```

`fixPrankNumber([0, 100, 200, 300, 150, 500])` は `[0, 100, 200, 300, 400, 500]` を返すべきです。

```js
assert.deepEqual(fixPrankNumber([0, 100, 200, 300, 150, 500]), [0, 100, 200, 300, 400, 500]);
```

`fixPrankNumber([400, 425, 400, 375, 350, 325, 300])` は `[450, 425, 400, 375, 350, 325, 300]` を返すべきです。

```js
assert.deepEqual(fixPrankNumber([400, 425, 400, 375, 350, 325, 300]), [450, 425, 400, 375, 350, 325, 300]);
```

`fixPrankNumber([-5, 5, 10, 15, 20])` は `[0, 5, 10, 15, 20]` を返すべきです。

```js
assert.deepEqual(fixPrankNumber([-5, 5, 10, 15, 20]), [0, 5, 10, 15, 20]);
```

# --seed--

## --seed-contents--

```js
function fixPrankNumber(arr) {

  return arr;
}
```

# --solutions--

```js
function fixPrankNumber(arr) {
  const diffs = arr.slice(1).map((v, i) => v - arr[i]);

  let step = diffs[0];
  for (let i = 0; i < diffs.length - 1; i++) {
    if (diffs[i] === diffs[i + 1]) {
      step = diffs[i];
      break;
    }
  }

  const result = [...arr];

  for (let i = 0; i < result.length - 1; i++) {
    if (result[i + 1] - result[i] !== step) {
      if (i === 0 && result[2] - result[1] === step) {
        result[0] = result[1] - step;
      } else {
        result[i + 1] = result[i] + step;
      }
      break;
    }
  }

  return result;
}
```
