---
id: 69b1028d6e265413d0198a2e
title: "チャレンジ 235: 大文字化されたフィボナッチ"
challengeType: 28
dashedName: challenge-235
---

# --description--

文字列が与えられたとき、インデックスがフィボナッチ数であればその文字を大文字にし、それ以外は小文字にした新しい文字列を返してください。

フィボナッチ数列は、各数字が直前の2つの数字の合計になる数の列です。数列の最初の10個の数字は `0`、`1`、`1`、`2`、`3`、`5`、`8`、`13`、`21`、`34` です。

- 最初の文字はインデックス `0` にあります。
- 文字でない文字のインデックスがフィボナッチ数であっても、その文字は変更しないでください。

# --hints--

`capitalizeFibonacci("hello world")` は `"HELLo woRld"` を返すべきです。

```js
assert.equal(capitalizeFibonacci("hello world"), "HELLo woRld");
```

`capitalizeFibonacci("HELLO WORLD")` は `"HELLo woRld"` を返すべきです。

```js
assert.equal(capitalizeFibonacci("HELLO WORLD"), "HELLo woRld");
```

`capitalizeFibonacci("hello, world!")` は `"HELLo, wOrld!"` を返すべきです。

```js
assert.equal(capitalizeFibonacci("hello, world!"), "HELLo, wOrld!");
```

`capitalizeFibonacci("The quick brown fox jumped over the lazy dog.")` は `"THE qUicK broWn fox jUmped over thE lazy dog."` を返すべきです。

```js
assert.equal(capitalizeFibonacci("The quick brown fox jumped over the lazy dog."), "THE qUicK broWn fox jUmped over thE lazy dog.");
```

`capitalizeFibonacci("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Proin pulvinar ex nibh, vel ullamcorper ligula egestas quis. Integer tincidunt fringilla accumsan. Integer et metus placerat, gravida felis at, pellentesque nisl.")` は `"LOREm ipSum dOlor sit amet, consecTetur adipiscing elit. proin pulvinar ex nibh, vel ullaMcorper ligula egestas quis. integer tincidunt fringillA accumsan. integer et metus placerat, gravida felis at, pellentesque nisl."` を返すべきです。

```js
assert.equal(capitalizeFibonacci("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Proin pulvinar ex nibh, vel ullamcorper ligula egestas quis. Integer tincidunt fringilla accumsan. Integer et metus placerat, gravida felis at, pellentesque nisl."), "LOREm ipSum dOlor sit amet, consecTetur adipiscing elit. proin pulvinar ex nibh, vel ullaMcorper ligula egestas quis. integer tincidunt fringillA accumsan. integer et metus placerat, gravida felis at, pellentesque nisl.");
```

# --seed--

## --seed-contents--

```js
function capitalizeFibonacci(str) {

  return str;
}
```

# --solutions--

```js
function capitalizeFibonacci(str) {
  const fibs = new Set([0, 1]);
  let a = 0, b = 1;
  while (b < str.length) {
    [a, b] = [b, a + b];
    fibs.add(b);
  }

  return str
    .split("")
    .map((char, i) => {
      if (!/[a-zA-Z]/.test(char)) return char;
      return fibs.has(i) ? char.toUpperCase() : char.toLowerCase();
    })
    .join("");
}
```
