---
id: 69b1028d6e265413d0198a2b
title: "チャレンジ 232: 期日"
challengeType: 28
dashedName: challenge-232
---

# --description--

日付の文字列が与えられたら、9か月後の日付を返してください。

- 与えられた文字列と返す文字列は `"YYYY-MM-DD"` の形式です。
- 9か月後の月に元の日付の数字が存在しない場合は、その月の最終日を返してください。

# --hints--

`getDueDate("2025-03-30")` は `"2025-12-30"` を返すべきです。

```js
assert.equal(getDueDate("2025-03-30"), "2025-12-30");
```

`getDueDate("2025-04-27")` は `"2026-01-27"` を返すべきです。

```js
assert.equal(getDueDate("2025-04-27"), "2026-01-27");
```

`getDueDate("2025-05-29")` は `"2026-02-28"` を返すべきです。

```js
assert.equal(getDueDate("2025-05-29"), "2026-02-28");
```

`getDueDate("2026-06-30")` は `"2027-03-30"` を返すべきです。

```js
assert.equal(getDueDate("2026-06-30"), "2027-03-30");
```

`getDueDate("2026-10-11")` は `"2027-07-11"` を返すべきです。

```js
assert.equal(getDueDate("2026-10-11"), "2027-07-11");
```

# --seed--

## --seed-contents--

```js
function getDueDate(dateStr) {

  return dateStr;
}
```

# --solutions--

```js
function getDueDate(dateStr) {
  let [year, month, day] = dateStr.split("-").map(Number);

  month += 9;

  year += Math.floor((month - 1) / 12);
  month = ((month - 1) % 12) + 1;

  const daysInMonth = new Date(year, month, 0).getDate();

  if (day > daysInMonth) {
    day = daysInMonth;
  }

  const mm = String(month).padStart(2, "0");
  const dd = String(day).padStart(2, "0");

  return `${year}-${mm}-${dd}`;
}
```
