---
id: 69b5b2be76ec8135a7fbe973
title: "チャレンジ 244: ルークとビショップの攻撃"
challengeType: 28
dashedName: challenge-244
---

# --description--

チェス盤上のルークの位置を示す文字列と、ビショップの位置を示す別の文字列が与えられたとき、どちらかの駒がもう一方を攻撃できるかどうかを判定してください。

標準的なチェス盤は8×8で、列は左から右へ`A`から`H`まで、行は下から上へ`1`から`8`までラベル付けされています。見た目は次のようになります：

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

- ルークは水平または垂直方向に好きなだけマスを移動できます。
- ビショップは任意の斜め方向に好きなだけマスを移動できます。
- ある駒がもう一方の駒の位置に移動できる場合、その駒はもう一方を攻撃できます。

返す値：

- ルークがビショップを攻撃できる場合は`"rook"`を返してください。
- ビショップがルークを攻撃できる場合は`"bishop"`を返してください。
- どちらの駒も互いを攻撃できない場合は`"neither"`を返してください。

# --hints--

`rookBishopAttack("A1", "A5")`は`"rook"`を返すべきです。

```js
assert.equal(rookBishopAttack("A1", "A5"), "rook");
```

`rookBishopAttack("C3", "F6")`は`"bishop"`を返すべきです。

```js
assert.equal(rookBishopAttack("C3", "F6"), "bishop");
```

`rookBishopAttack("D4", "D7")`は`"rook"`を返すべきです。

```js
assert.equal(rookBishopAttack("D4", "D7"), "rook");
```

`rookBishopAttack("B7", "H1")`は`"bishop"`を返すべきです。

```js
assert.equal(rookBishopAttack("B7", "H1"), "bishop");
```

`rookBishopAttack("B3", "C5")`は`"neither"`を返すべきです。

```js
assert.equal(rookBishopAttack("B3", "C5"), "neither");
```

`rookBishopAttack("G3", "E8")`は`"neither"`を返すべきです。

```js
assert.equal(rookBishopAttack("G3", "E8"), "neither");
```

# --seed--

## --seed-contents--

```js
function rookBishopAttack(rook, bishop) {

  return rook;
}
```

# --solutions--

```js
function rookBishopAttack(rook, bishop) {
  const colDiff = Math.abs(rook.charCodeAt(0) - bishop.charCodeAt(0));
  const rowDiff = Math.abs(rook[1] - bishop[1]);

  if (rook[0] === bishop[0] || rook[1] === bishop[1]) return "rook";
  if (colDiff === rowDiff) return "bishop";
  return "neither";
}
```
