---
id: 69b559d2903b9e4afe9075f7
title: "チャレンジ 237: 方程式のバリデーション"
challengeType: 28
dashedName: challenge-237
---

# --description--

数学の方程式を表す文字列が与えられたとき、それが正しいかどうかを判定してください。

- 左辺には最大で3つの正の整数と演算子 `+`、`-`、`*`、`/` が含まれる場合があります。
- 方程式は `"number operator number = number"` の形式で与えられます（左辺に2つまたは3つの数字がある場合）。例えば `"2 + 2 = 4"` や `"2 + 3 - 1 = 4"` のようになります。
- 右辺は常に単一の整数です。

標準の演算順序に従ってください：掛け算と割り算は足し算と引き算より先に、左から右へ評価されます。

# --hints--

`isValidEquation("2 + 2 = 4")` は `true` を返すべきです。

```js
assert.isTrue(isValidEquation("2 + 2 = 4"));
```

`isValidEquation("2 + 3 - 1 = 4")` は `true` を返すべきです。

```js
assert.isTrue(isValidEquation("2 + 3 - 1 = 4"));
```

`isValidEquation("8 / 2 = 4")` は `true` を返すべきです。

```js
assert.isTrue(isValidEquation("8 / 2 = 4"));
```

`isValidEquation("10 * 5 = 50")` は `true` を返すべきです。

```js
assert.isTrue(isValidEquation("10 * 5 = 50"));
```

`isValidEquation("2 - 2 = 0")` は `true` を返すべきです。

```js
assert.isTrue(isValidEquation("2 - 2 = 0"));
```

`isValidEquation("2 + 9 / 3 = 5")` は `true` を返すべきです。

```js
assert.isTrue(isValidEquation("2 + 9 / 3 = 5"));
```

`isValidEquation("20 - 2 * 3 = 14")` は `true` を返すべきです。

```js
assert.isTrue(isValidEquation("20 - 2 * 3 = 14"));
```

`isValidEquation("2 + 5 = 6")` は `false` を返すべきです。

```js
assert.isFalse(isValidEquation("2 + 5 = 6"));
```

`isValidEquation("10 - 2 * 3 = 24")` は `false` を返すべきです。

```js
assert.isFalse(isValidEquation("10 - 2 * 3 = 24"));
```

`isValidEquation("3 + 9 / 3 = 4")` は `false` を返すべきです。

```js
assert.isFalse(isValidEquation("3 + 9 / 3 = 4"));
```

# --seed--

## --seed-contents--

```js
function isValidEquation(equation) {

  return equation;
}
```

# --solutions--

```js
function isValidEquation(equation) {
  const [left, right] = equation.split(" = ");
  const tokens = left.split(" ");

  for (let i = 1; i < tokens.length - 1; i += 2) {
    const operator = tokens[i];
    if (operator === "*" || operator === "/") {
      const leftOperand = Number(tokens[i - 1]);
      const rightOperand = Number(tokens[i + 1]);
      const computed = operator === "*" ? leftOperand * rightOperand : leftOperand / rightOperand;
      tokens.splice(i - 1, 3, String(computed));
      i -= 2;
    }
  }

  let result = Number(tokens[0]);
  for (let i = 1; i < tokens.length - 1; i += 2) {
    const operator = tokens[i];
    const operand = Number(tokens[i + 1]);
    result = operator === "+" ? result + operand : result - operand;
  }

  return result === Number(right);
}
```
