---
id: 69b559d2903b9e4afe9075f8
title: "チャレンジ 238: 桁の回転エスケープ"
challengeType: 28
dashedName: challenge-238
---

# --description--

正の整数が与えられたとき、その数自身またはその回転のいずれかが、その桁数で割り切れるかどうかを判定してください。

回転とは、最初の桁を末尾に移動させることを意味します。例えば、1回回転させると、123は231になります。

- まず回転なしの`0`（与えられた数）をチェックしてください。
- 与えられる数にはゼロは含まれません。
- 割り切れる回転数が見つかった場合はその数を返し、なければ`"none"`を返してください。

# --hints--

`getRotation(123)`は`0`を返すべきです。

```js
assert.equal(getRotation(123), 0);
```

`getRotation(13579)`は`3`を返すべきです。

```js
assert.equal(getRotation(13579), 3);
```

`getRotation(24681)`は`"none"`を返すべきです。

```js
assert.equal(getRotation(24681), "none");
```

`getRotation(84138789345)`は`6`を返すべきです。

```js
assert.equal(getRotation(84138789345), 6);
```

# --seed--

## --seed-contents--

```js
function getRotation(n) {

  return n;
}
```

# --solutions--

```js
function getRotation(n) {
  const str = String(n);
  const digitCount = str.length;
  let current = str;

  for (let i = 0; i < digitCount; i++) {
    if (Number(current) % digitCount === 0) return i;
    current = current.slice(1) + current[0];
  }

  return "none";
}
```
