---
id: 69b58ce40693f140c84c8559
title: "チャレンジ 240: 回文の文字"
challengeType: 28
dashedName: challenge-240
---

# --description--

文字列が与えられたとき、それが回文かどうかを判定し、長さが奇数なら中央の文字を、偶数なら中央の2文字を返してください。

- 回文とは、前から読んでも後ろから読んでも同じ文字列のことです。
- 回文でない場合は `"none"` を返してください。

# --hints--

`palindromeLocator("racecar")` は `"e"` を返すべきです。

```js
assert.equal(palindromeLocator("racecar"), "e");
```

`palindromeLocator("level")` は `"v"` を返すべきです。

```js
assert.equal(palindromeLocator("level"), "v");
```

`palindromeLocator("freecodecamp")` は `"none"` を返すべきです。

```js
assert.equal(palindromeLocator("freecodecamp"), "none");
```

`palindromeLocator("noon")` は `"oo"` を返すべきです。

```js
assert.equal(palindromeLocator("noon"), "oo");
```

`palindromeLocator("11100111")` は `"00"` を返すべきです。

```js
assert.equal(palindromeLocator("11100111"), "00");
```

# --seed--

## --seed-contents--

```js
function palindromeLocator(str) {

  return str;
}
```

# --solutions--

```js
function palindromeLocator(str) {
  if (str !== str.split("").reverse().join("")) return "none";

  const mid = Math.floor(str.length / 2);
  return str.length % 2 === 1 ? str[mid] : str[mid - 1] + str[mid];
}
```
