---
id: 69b5b2be76ec8135a7fbe974
title: "チャレンジ 245: スパイラルマトリックス"
challengeType: 28
dashedName: challenge-245
---

# --description--

2次元マトリックスが与えられたとき、すべての値を時計回りの順序で並べた平坦な配列を返してください。

返される配列は、左上の値から始まり、上の行を右に進み、次に右の列を下に進み、次に下の行を左に進み、最後に左の列を上に進みます。残りの層があれば、内側に向かって繰り返します。

例えば、次のような場合：

```json
[
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9]
]
```

`[1, 2, 3, 6, 9, 8, 7, 4, 5]` を返してください。

# --hints--

`spiralMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]])` は `[1, 2, 3, 6, 9, 8, 7, 4, 5]` を返すべきです。

```js
assert.deepEqual(spiralMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]]), [1, 2, 3, 6, 9, 8, 7, 4, 5]);
```

`spiralMatrix([["a", "b", "c", "d"], ["l", "m", "n", "e"], ["k", "p", "o", "f"], ["j", "i", "h", "g"]])` は `["a", "b", "c", "d", "e", "f", "g", "h", "i", "j", "k", "l", "m", "n", "o", "p"]` を返すべきです。

```js
assert.deepEqual(spiralMatrix([["a", "b", "c", "d"], ["l", "m", "n", "e"], ["k", "p", "o", "f"], ["j", "i", "h", "g"]]), ["a", "b", "c", "d", "e", "f", "g", "h", "i", "j", "k", "l", "m", "n", "o", "p"]);
```

`spiralMatrix([[true, false, false], [false, true, true], [false, true, false], [true, true, false]])` は `[true, false, false, true, false, false, true, true, false, false, true, true]` を返すべきです。

```js
assert.deepEqual(spiralMatrix([[true, false, false], [false, true, true], [false, true, false], [true, true, false]]), [true, false, false, true, false, false, true, true, false, false, true, true]);
```

`spiralMatrix([[25, 24, 23, 22, 21], [10, 9, 8, 7, 20], [11, 2, 1, 6, 19], [12, 3, 4, 5, 18], [13, 14, 15, 16, 17]])` は `[25, 24, 23, 22, 21, 20, 19, 18, 17, 16, 15, 14, 13, 12, 11, 10, 9, 8, 7, 6, 5, 4, 3, 2, 1]` を返すべきです。

```js
assert.deepEqual(spiralMatrix([[25, 24, 23, 22, 21], [10, 9, 8, 7, 20], [11, 2, 1, 6, 19], [12, 3, 4, 5, 18], [13, 14, 15, 16, 17]]), [25, 24, 23, 22, 21, 20, 19, 18, 17, 16, 15, 14, 13, 12, 11, 10, 9, 8, 7, 6, 5, 4, 3, 2, 1]);
```

# --seed--

## --seed-contents--

```js
function spiralMatrix(matrix) {

    return matrix
}
```

# --solutions--

```js
function spiralMatrix(matrix) {
  const result = [];
  let top = 0, bottom = matrix.length - 1;
  let left = 0, right = matrix[0].length - 1;

  while (top <= bottom && left <= right) {
    for (let i = left; i <= right; i++) result.push(matrix[top][i]);
    top++;

    for (let i = top; i <= bottom; i++) result.push(matrix[i][right]);
    right--;

    if (top <= bottom) {
      for (let i = right; i >= left; i--) result.push(matrix[bottom][i]);
      bottom--;
    }

    if (left <= right) {
      for (let i = bottom; i >= top; i--) result.push(matrix[i][left]);
      left++;
    }
  }

  return result;
}
```
