---
id: 69b559d2903b9e4afe9075f9
title: "チャレンジ 239: 今日は何曜日？"
challengeType: 28
dashedName: challenge-239
---

# --description--

ミリ秒単位のUnixタイムスタンプが与えられたとき、曜日を返してください。

有効な返却曜日は以下の通りです：

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

タイムゾーンは無視してください。

# --hints--

`getDayOfWeek(1775492249000)` は `"Monday"` を返すべきです。

```js
assert.equal(getDayOfWeek(1775492249000), "Monday");
```

`getDayOfWeek(1766246400000)` は `"Saturday"` を返すべきです。

```js
assert.equal(getDayOfWeek(1766246400000), "Saturday");
```

`getDayOfWeek(33791256000000)` は `"Tuesday"` を返すべきです。

```js
assert.equal(getDayOfWeek(33791256000000), "Tuesday");
```

`getDayOfWeek(1773576000000)` は `"Sunday"` を返すべきです。

```js
assert.equal(getDayOfWeek(1773576000000), "Sunday");
```

`getDayOfWeek(0)` は `"Thursday"` を返すべきです。

```js
assert.equal(getDayOfWeek(0), "Thursday");
```

# --seed--

## --seed-contents--

```js
function getDayOfWeek(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function getDayOfWeek(timestamp) {
  return new Date(timestamp).toLocaleDateString("en-US", { weekday: "long", timeZone: "UTC" });
}
```
