---
id: 69b58ce40693f140c84c855c
title: "チャレンジ 243: ルークの攻撃"
challengeType: 28
dashedName: challenge-243
---

# --description--

チェス盤上の2つのルークの位置を表す2つの文字列が与えられたとき、それらが互いに攻撃できるかどうかを判定してください。

標準的なチェス盤は8×8で、列は左から右に`A`から`H`まで、行は下から上に`1`から`8`までラベル付けされています。見た目は次のようになります。

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

ルークは水平または垂直方向に好きなだけマスを移動できます。したがって、同じ行または列にいる場合は互いに攻撃できます。

# --hints--

`rookAttack("A1", "A8")`は`true`を返すべきです。

```js
assert.isTrue(rookAttack("A1", "A8"));
```

`rookAttack("B4", "F4")`は`true`を返すべきです。

```js
assert.isTrue(rookAttack("B4", "F4"));
```

`rookAttack("E3", "D4")`は`false`を返すべきです。

```js
assert.isFalse(rookAttack("E3", "D4"));
```

`rookAttack("H7", "F6")`は`false`を返すべきです。

```js
assert.isFalse(rookAttack("H7", "F6"));
```

# --seed--

## --seed-contents--

```js
function rookAttack(rook1, rook2) {

  return rook1;
}
```

# --solutions--

```js
function rookAttack(rook1, rook2) {
  return rook1[0] === rook2[0] || rook1[1] === rook2[1];
}
```
