---
id: 69b58ce40693f140c84c855b
title: "チャレンジ 242: 次のビンゴ番号"
challengeType: 28
dashedName: challenge-242
---

# --description--

ビンゴ番号が与えられたら、次のビンゴ番号を順番に返してください。

ビンゴ番号は、次のチャートに従った範囲の数字が続く1文字のアルファベットです。

| Letter | Number Range |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

例えば、`"B10"` が与えられたら、次のビンゴ番号である `"B11"` を返してください。最後のビンゴ番号が与えられた場合は、`"B1"` を返してください。

# --hints--

`getNextBingoNumber("B10")` は `"B11"` を返すべきです。

```js
assert.equal(getNextBingoNumber("B10"), "B11");
```

`getNextBingoNumber("N33")` は `"N34"` を返すべきです。

```js
assert.equal(getNextBingoNumber("N33"), "N34");
```

`getNextBingoNumber("I30")` は `"N31"` を返すべきです。

```js
assert.equal(getNextBingoNumber("I30"), "N31");
```

`getNextBingoNumber("G60")` は `"O61"` を返すべきです。

```js
assert.equal(getNextBingoNumber("G60"), "O61");
```

`getNextBingoNumber("O75")` は `"B1"` を返すべきです。

```js
assert.equal(getNextBingoNumber("O75"), "B1");
```

# --seed--

## --seed-contents--

```js
function getNextBingoNumber(n) {

  return n;
}
```

# --solutions--

```js
function getNextBingoNumber(n) {
  const num = parseInt(n.slice(1));
  const next = num === 75 ? 1 : num + 1;

  if (next <= 15) return "B" + next;
  if (next <= 30) return "I" + next;
  if (next <= 45) return "N" + next;
  if (next <= 60) return "G" + next;
  return "O" + next;
}
```
