---
id: 69bc6cb30c1d112a2e110a04
title: "チャレンジ247：最後の文字"
challengeType: 28
dashedName: challenge-247
---

# --description--

文字列が与えられたとき、その文字列の中でアルファベット順で最後に現れる文字を返してください。

- もし複数の文字がアルファベット順で最後に並んでいる場合は、最初に現れた文字を返してください。
- 文字以外のすべての文字は無視してください。

# --hints--

`getLastLetter("world")` は `"w"` を返すべきです。

```js
assert.equal(getLastLetter("world"), "w");
```

`getLastLetter("Hello World")` は `"W"` を返すべきです。

```js
assert.equal(getLastLetter("Hello World"), "W");
```

`getLastLetter("The quick brown fox jumped over the lazy dog.")` は `"z"` を返すべきです。

```js
assert.equal(getLastLetter("The quick brown fox jumped over the lazy dog."), "z");
```

`getLastLetter("HeLl0")` は `"L"` を返すべきです。

```js
assert.equal(getLastLetter("HeLl0"), "L");
```

`getLastLetter("!#$ er@R asd fT.,> 2t0e9")` は `"T"` を返すべきです。

```js
assert.equal(getLastLetter("!#$ er@R asd fT.,> 2t0e9"), "T");
```

# --seed--

## --seed-contents--

```js
function getLastLetter(str) {

  return str;
}
```

# --solutions--

```js
function getLastLetter(str) {
  const letters = str.split('').filter(char => /[a-zA-Z]/.test(char));
  return letters.reduce((last, char) => char.toLowerCase() > last.toLowerCase() ? char : last);
}
```
