---
id: 69bc6cb30c1d112a2e110a08
title: "チャレンジ 251: 配列の合計を見つける"
challengeType: 28
dashedName: challenge-251
---

# --description--

数値の配列と目標の数が与えられたとき、目標に合計される2つ以上の最初の部分集合を返してください。

- 「最初の」部分集合とは、要素のインデックスが可能な限り低く、最も早いインデックスを優先するものです。
- 配列内の各数値は一度だけ使用できます。
- 有効な部分集合が存在しない場合は `"Sum not found"` を返してください。

一致する数値は、元の配列に現れる順序で配列として返してください。

# --hints--

`findSum([1, 3, 5, 7], 6)` は `[1, 5]` を返すべきです。

```js
assert.deepEqual(findSum([1, 3, 5, 7], 6), [1, 5]);
```

`findSum([1, 2, 3, 4, 5], 5)` は `[1, 4]` を返すべきです。

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 5), [1, 4]);
```

`findSum([1, 2, 3, 4, 5], 6)` は `[1, 2, 3]` を返すべきです。

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 6), [1, 2, 3]);
```

`findSum([-1, -2, 3, 4], 1)` は `[-1, -2, 4]` を返すべきです。

```js
assert.deepEqual(findSum([-1, -2, 3, 4], 1), [-1, -2, 4]);
```

`findSum([3, 1, 4, 1, 5, 9, 2, 6], 10)` は `[3, 1, 4, 2]` を返すべきです。

```js
assert.deepEqual(findSum([3, 1, 4, 1, 5, 9, 2, 6], 10), [3, 1, 4, 2]);
```

`findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20)` は `[1, 2, 3, 5, 9]` を返すべきです。

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20), [1, 2, 3, 5, 9]);
```

`findSum([7, 9, 4, 2, 5], 10)` は `"Sum not found"` を返すべきです。

```js
assert.equal(findSum([7, 9, 4, 2, 5], 10), "Sum not found");
```

# --seed--

## --seed-contents--

```js
function findSum(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findSum(arr, target) {
  function search(start, remaining, current) {
    if (remaining === 0 && current.length >= 2) return current;
    if (start >= arr.length) return null;

    for (let i = start; i < arr.length; i++) {
      const result = search(i + 1, remaining - arr[i], [...current, arr[i]]);
      if (result) return result;
    }

    return null;
  }

  return search(0, target, []) ?? "Sum not found";
}
```
