---
id: 69bc6cb30c1d112a2e110a09
title: "チャレンジ252：ユニークな階段の登り方"
challengeType: 28
dashedName: challenge-252
---

# --description--

階段の段数が与えられたとき、1回に1段または2段ずつ登る方法で、何通りの異なる登り方があるかを返してください。

# --hints--

`getUniqueClimbs(4)` は `5` を返すべきです。

```js
assert.equal(getUniqueClimbs(4), 5);
```

`getUniqueClimbs(5)` は `8` を返すべきです。

```js
assert.equal(getUniqueClimbs(5), 8);
```

`getUniqueClimbs(10)` は `89` を返すべきです。

```js
assert.equal(getUniqueClimbs(10), 89);
```

`getUniqueClimbs(18)` は `4181` を返すべきです。

```js
assert.equal(getUniqueClimbs(18), 4181);
```

`getUniqueClimbs(29)` は `832040` を返すべきです。

```js
assert.equal(getUniqueClimbs(29), 832040);
```

`getUniqueClimbs(50)` は `20365011074` を返すべきです。

```js
assert.equal(getUniqueClimbs(50), 20365011074);
```

# --seed--

## --seed-contents--

```js
function getUniqueClimbs(steps) {

  return steps;
}
```

# --solutions--

```js
function getUniqueClimbs(steps) {
  if (steps <= 0) return 0;
  if (steps === 1) return 1;
  if (steps === 2) return 2;

  let prev2 = 1, prev1 = 2;
  for (let i = 3; i <= steps; i++) {
    [prev2, prev1] = [prev1, prev2 + prev1];
  }
  return prev1;
}
```
