---
id: 69bc6cb30c1d112a2e110a06
title: "チャレンジ249：文字列の計算"
challengeType: 28
dashedName: challenge-249
---

# --description--

数字とその他の文字が混ざった文字列が与えられたとき、数字の間にある数字以外の文字の数に基づいて数字同士の計算を行ってください。

- 2つの数字を区切る文字の数が偶数なら、足し算をします。
- 奇数なら引き算をします。
- 連続する数字は1つの数字として扱います。
- 計算は左から右へ順に行います。
- 先頭と末尾の数字以外の文字は無視します。

例えば、`"3ab10c8"`が与えられたら、`5`を返してください。3と10の間に偶数の文字があるので足して13になります。次に、13と8の間に奇数の文字があるので13から8を引きます。

# --hints--

`doMath("3ab10c8")`は`5`を返すべきです。

```js
assert.equal(doMath("3ab10c8"), 5);
```

`doMath("6MINUS4")`は`2`を返すべきです。

```js
assert.equal(doMath("6MINUS4"), 2);
```

`doMath("9plus3")`は`12`を返すべきです。

```js
assert.equal(doMath("9plus3"), 12);
```

`doMath("5fkwo#10i#%.<>15P=@20!#B/25")`は`15`を返すべきです。

```js
assert.equal(doMath("5fkwo#10i#%.<>15P=@20!#B/25"), 15);
```

`doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt")`は`67`を返すべきです。

```js
assert.equal(doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt"), 67);
```

# --seed--

## --seed-contents--

```js
function doMath(str) {

  return str;
}
```

# --solutions--

```js
function doMath(str) {
  const tokens = str.match(/(\d+)|(\D+)/g);
  let result = null;
  let pendingOp = null;

  for (const token of tokens) {
    if (/^\d+$/.test(token)) {
      if (result === null) {
        result = parseInt(token);
      } else {
        result = pendingOp === 'add' ? result + parseInt(token) : result - parseInt(token);
      }
    } else {
      pendingOp = token.length % 2 === 0 ? 'add' : 'subtract';
    }
  }

  return result;
}
```
