---
id: 69bc6cb30c1d112a2e110a03
title: "チャレンジ 246: 名前のイニシャル"
challengeType: 28
dashedName: challenge-246
---

# --description--

フルネームが文字列で与えられたとき、そのイニシャルを返してください。

- イニシャルを取る名前はスペースで区切られています。
- イニシャルは大文字にしてください。
- イニシャルはドットで区切ってください。

例えば、`"Tommy Millwood"` は `"T.M."` を返します。

# --hints--

`getInitials("Tommy Millwood")` は `"T.M."` を返すはずです。

```js
assert.equal(getInitials("Tommy Millwood"), "T.M.");
```

`getInitials("Savanna Puddlesplash")` は `"S.P."` を返すはずです。

```js
assert.equal(getInitials("Savanna Puddlesplash"), "S.P.");
```

`getInitials("Frances Cowell Conrad")` は `"F.C.C."` を返すはずです。

```js
assert.equal(getInitials("Frances Cowell Conrad"), "F.C.C.");
```

`getInitials("Dragon")` は `"D."` を返すはずです。

```js
assert.equal(getInitials("Dragon"), "D.");
```

`getInitials("Dorothy Vera Clump Haverstock Norris")` は `"D.V.C.H.N."` を返すはずです。

```js
assert.equal(getInitials("Dorothy Vera Clump Haverstock Norris"), "D.V.C.H.N.");
```

# --seed--

## --seed-contents--

```js
function getInitials(name) {

  return name;
}
```

# --solutions--

```js
function getInitials(name) {
  return name.split(' ').map(word => word[0].toUpperCase() + '.').join('');
}
```
