---
id: 69c5f3d787b1725d5f00c8bb
title: "チャレンジ 257: 文字列圧縮器"
challengeType: 28
dashedName: challenge-257
---

# --description--

文字列が与えられたとき、次のルールに従って圧縮した文字列を返してください。

- 単語の最初の出現はそのままにします。
- それ以降の出現は、最初の出現の位置で置き換えます。最初の単語は位置1とします。
- 単語は1つのスペースで区切られます。

例えば、`"practice makes perfect and perfect practice makes perfect"`が与えられたら、`"practice makes perfect and 3 1 2 3"`を返してください。

# --hints--

`compress("practice makes perfect and perfect practice makes perfect")`は`"practice makes perfect and 3 1 2 3"`を返すべきです。

```js
assert.equal(compress("practice makes perfect and perfect practice makes perfect"), "practice makes perfect and 3 1 2 3");
```

`compress("hello hello hello")`は`"hello 1 1"`を返すべきです。

```js
assert.equal(compress("hello hello hello"), "hello 1 1");
```

`compress("the cat sat on the mat on which the cat sat")`は`"the cat sat on 1 mat 4 which 1 2 3"`を返すべきです。

```js
assert.equal(compress("the cat sat on the mat on which the cat sat"), "the cat sat on 1 mat 4 which 1 2 3");
```

`compress("the more you know the more you realize you don't know")`は`"the more you know 1 2 3 realize 3 don't 4"`を返すべきです。

```js
assert.equal(compress("the more you know the more you realize you don't know"), "the more you know 1 2 3 realize 3 don't 4");
```

`compress("lorem ipsum dolor sit per elit donec sit nostra libero per donec ligula sit gravida at elit vitae a elit sodales donec en donec at dolor nam ligula dignissim risus at ligula per nam ipsum ipsum gravida en elit per ipsum ligula en gravida per sodales sit at nam lorem sit per libero en ipsum elit sit sodales sit risus elit risus ipsum elit at gravida vitae en dignissim nam sit vitae sollicitudin per nostra per sit libero")`は`"lorem ipsum dolor sit per elit donec 4 nostra libero 5 7 ligula 4 gravida at 6 vitae a 6 sodales 7 en 7 16 3 nam 13 dignissim risus 16 13 5 27 2 2 15 23 6 5 2 13 23 15 5 21 4 16 27 1 4 5 10 23 2 6 4 21 4 30 6 30 2 6 16 15 18 23 29 27 4 18 sollicitudin 5 9 5 4 10"`を返すべきです。

```js
assert.equal(compress("lorem ipsum dolor sit per elit donec sit nostra libero per donec ligula sit gravida at elit vitae a elit sodales donec en donec at dolor nam ligula dignissim risus at ligula per nam ipsum ipsum gravida en elit per ipsum ligula en gravida per sodales sit at nam lorem sit per libero en ipsum elit sit sodales sit risus elit risus ipsum elit at gravida vitae en dignissim nam sit vitae sollicitudin per nostra per sit libero"), "lorem ipsum dolor sit per elit donec 4 nostra libero 5 7 ligula 4 gravida at 6 vitae a 6 sodales 7 en 7 16 3 nam 13 dignissim risus 16 13 5 27 2 2 15 23 6 5 2 13 23 15 5 21 4 16 27 1 4 5 10 23 2 6 4 21 4 30 6 30 2 6 16 15 18 23 29 27 4 18 sollicitudin 5 9 5 4 10");
```

# --seed--

## --seed-contents--

```js
function compress(str) {

  return str;
}
```

# --solutions--

```js
function compress(str) {
  const words = str.split(" ");
  const seen = {};

  return words.map((word, i) => {
    if (seen[word] !== undefined) {
      return seen[word];
    }
    seen[word] = i + 1;
    return word;
  }).join(" ");
}
```
