---
id: 69c5f3d787b1725d5f00c8ba
title: "チャレンジ 256: 最も近い時間の方向"
challengeType: 28
dashedName: challenge-256
---

# --description--

2つの時間が与えられたとき、最初の時間から2番目の時間へ進む方向と戻る方向のどちらが速いかを判定してください。

- 時間は24時間表記で与えられます (`"HH:MM"`)
- 時計は一周します（進む場合は23:59の次は00:00、戻る場合は00:00の前は23:59）

返す値：

- 進む方向が短ければ `"forward"`
- 戻る方向が短ければ `"backward"`
- 両方向とも同じ時間なら `"equal"`

# --hints--

`getDirection("10:00", "12:00")` は `"forward"` を返すべきです。

```js
assert.equal(getDirection("10:00", "12:00"), "forward");
```

`getDirection("11:00", "05:00")` は `"backward"` を返すべきです。

```js
assert.equal(getDirection("11:00", "05:00"), "backward");
```

`getDirection("00:00", "12:00")` は `"equal"` を返すべきです。

```js
assert.equal(getDirection("00:00", "12:00"), "equal");
```

`getDirection("15:45", "01:10")` は `"forward"` を返すべきです。

```js
assert.equal(getDirection("15:45", "01:10"), "forward");
```

`getDirection("03:30", "19:50")` は `"backward"` を返すべきです。

```js
assert.equal(getDirection("03:30", "19:50"), "backward");
```

`getDirection("06:30", "18:30")` は `"equal"` を返すべきです。

```js
assert.equal(getDirection("06:30", "18:30"), "equal");
```

# --seed--

## --seed-contents--

```js
function getDirection(time1, time2) {

  return time1;
}
```

# --solutions--

```js
function getDirection(time1, time2) {
  function toMinutes(t) {
    const [h, m] = t.split(":").map(Number);
    return h * 60 + m;
  }

  const start = toMinutes(time1);
  const end = toMinutes(time2);

  const forward = (end - start + 1440) % 1440;
  const backward = (start - end + 1440) % 1440;

  if (forward < backward) return "forward";
  if (backward < forward) return "backward";
  return "equal";
}
```
