---
id: 69c5f3d787b1725d5f00c8bd
title: "チャレンジ 259: FizzBuzz エクスプロージョン"
challengeType: 28
dashedName: challenge-259
---

# --description--

整数が与えられたとき、次のルールを使って単語 `"fizzbuzz"` を少なくとも指定された数の `"z"` を含む文字列に変えるまでにかかるステップ数を返してください。

- 文字列 `"fizzbuzz"` から始めます。
- 各ステップで、文字列内の文字の位置（最初の `"f"` は位置1）を使って標準のFizzBuzzルールを適用します。
  - 文字の位置が3で割り切れる場合、その文字を `"fizz"` に置き換えます。
  - 文字の位置が5で割り切れる場合、その文字を `"buzz"` に置き換えます。
  - 文字の位置が3と5の両方で割り切れる場合、その文字を `"fizzbuzz"` に置き換えます。

つまり、1ステップ後に `"fizzbuzz"` は `"fifizzzbuzzfizzzz"` に変わり、これは9つの `"z"` を含みます。

# --hints--

`explodeFizzbuzz(9)` は `1` を返すべきです。

```js
assert.equal(explodeFizzbuzz(9), 1);
```

`explodeFizzbuzz(15)` は `2` を返すべきです。

```js
assert.equal(explodeFizzbuzz(15), 2);
```

`explodeFizzbuzz(51)` は `3` を返すべきです。

```js
assert.equal(explodeFizzbuzz(51), 3);
```

`explodeFizzbuzz(52)` は `4` を返すべきです。

```js
assert.equal(explodeFizzbuzz(52), 4);
```

`explodeFizzbuzz(359)` は `5` を返すべきです。

```js
assert.equal(explodeFizzbuzz(359), 5);
```

`explodeFizzbuzz(789)` は `6` を返すべきです。

```js
assert.equal(explodeFizzbuzz(789), 6);
```

`explodeFizzbuzz(54482)` は `11` を返すべきです。

```js
assert.equal(explodeFizzbuzz(54482), 11);
```

`explodeFizzbuzz(1000000)` は `14` を返すべきです。

```js
assert.equal(explodeFizzbuzz(1000000), 14);
```

# --seed--

## --seed-contents--

```js
function explodeFizzbuzz(targetZCount) {

  return targetZCount;
}
```

# --solutions--

```js
function explodeFizzbuzz(targetZCount) {
  let str = "fizzbuzz";
  let steps = 0;

  while ((str.match(/z/g) || []).length < targetZCount) {
    let next = "";
    for (let i = 0; i < str.length; i++) {
      const pos = i + 1;
      if (pos % 15 === 0) next += "fizzbuzz";
      else if (pos % 3 === 0) next += "fizz";
      else if (pos % 5 === 0) next += "buzz";
      else next += str[i];
    }
    str = next;
    steps++;
  }

  return steps;
}
```
