---
id: 69c5f3d787b1725d5f00c8b7
title: "チャレンジ 253: 頭字語ファインダー"
challengeType: 28
dashedName: challenge-253
---

# --description--

頭字語を表す文字列が与えられたとき、以下のリストからその頭字語が属する組織の正式名称を返してください。

- `"National Avocado Storage Authority"`
- `"Cats Infiltration Agency"`
- `"Fluffy Beanbag Inspectors"`
- `"Department Of Jelly"`
- `"Wild Honey Organization"`
- `"Eating Pancakes Administration"`

与えられた頭字語の各文字は、その組織の各単語の最初の文字と同じ順序で一致する必要があります。

# --hints--

`findOrg("NASA")` は `"National Avocado Storage Authority"` を返すべきです。

```js
assert.equal(findOrg("NASA"), "National Avocado Storage Authority");
```

`findOrg("CIA")` は `"Cats Infiltration Agency"` を返すべきです。

```js
assert.equal(findOrg("CIA"), "Cats Infiltration Agency");
```

`findOrg("FBI")` は `"Fluffy Beanbag Inspectors"` を返すべきです。

```js
assert.equal(findOrg("FBI"), "Fluffy Beanbag Inspectors");
```

`findOrg("DOJ")` は `"Department Of Jelly"` を返すべきです。

```js
assert.equal(findOrg("DOJ"), "Department Of Jelly");
```

`findOrg("WHO")` は `"Wild Honey Organization"` を返すべきです。

```js
assert.equal(findOrg("WHO"), "Wild Honey Organization");
```

`findOrg("EPA")` は `"Eating Pancakes Administration"` を返すべきです。

```js
assert.equal(findOrg("EPA"), "Eating Pancakes Administration");
```

# --seed--

## --seed-contents--

```js
function findOrg(acronym) {

  return acronym;
}
```

# --solutions--

```js
function findOrg(acronym) {
  const orgs = [
    "National Avocado Storage Authority",
    "Cats Infiltration Agency",
    "Fluffy Beanbag Inspectors",
    "Department Of Jelly",
    "Wild Honey Organization",
    "Eating Pancakes Administration"
  ];

  for (let org of orgs) {
    const initials = org
      .split(" ")
      .map(word => word[0])
      .join("");

    if (initials.toUpperCase() === acronym.toUpperCase()) {
      return org;
    }
  }

  return "not found";
}
```
