---
id: 69c5f3d787b1725d5f00c8b8
title: "チャレンジ 254: 奇数文字の単語"
challengeType: 28
dashedName: challenge-254
---

# --description--

文字列で与えられた単語のうち、文字数が奇数の単語だけを返してください。

- 与えられた文字列の単語は、単一のスペースで区切られています。
- 単語は単一のスペースで区切って返してください。

# --hints--

`getOddWords("This is a super good test")` は `"a super"` を返すべきです。

```js
assert.equal(getOddWords("This is a super good test"), "a super");
```

`getOddWords("one two three four")` は `"one two three"` を返すべきです。

```js
assert.equal(getOddWords("one two three four"), "one two three");
```

`getOddWords("banana split sundae with rainbow sprinkles on top")` は `"split rainbow sprinkles top"` を返すべきです。

```js
assert.equal(getOddWords("banana split sundae with rainbow sprinkles on top"), "split rainbow sprinkles top");
```

`getOddWords("The quick brown fox jumped over the lazy river")` は `"The quick brown fox the river"` を返すべきです。

```js
assert.equal(getOddWords("The quick brown fox jumped over the lazy river"), "The quick brown fox the river");
```

# --seed--

## --seed-contents--

```js
function getOddWords(str) {

  return str;
}
```

# --solutions--

```js
function getOddWords(str) {
  return str
    .split(" ")
    .filter(word => word.length % 2 === 1)
    .join(" ");
}
```
