---
id: 69c5f3d787b1725d5f00c8b9
title: "チャレンジ 255: アースデイ クリーンアップクルー"
challengeType: 28
dashedName: challenge-255
---

# --description--

今日はアースデイです。清掃したアイテムの配列が与えられたとき、以下のルールに基づいて合計のクリーンアップスコアを返してください。

与えられるアイテムは以下のいずれかです：

| アイテム | 基本値 |
| - | - |
| `"bottle"` | 10 |
| `"can"` | 6 |
| `"bag"` | 8 |
| `"tire"` | 35 |
| `"straw"` | 4 |
| `"cardboard"` | 3 |
| `"newspaper"` | 3 |
| `"shoe"` | 12 |
| `"electronics"` | 25 |
| `"battery"` | 18 |
| `"mattress"` | 38 |

レアアイテムは `["rare", value]` で表されます。例えば `["rare", 80]` です。レアアイテムには連続ボーナスは付きません。

- 連続ボーナス：同じアイテムが連続して現れると、ボーナスポイントが増加します。
  - 1回目の連続出現：基本値
  - 2回目：基本値 + 1
  - 3回目：基本値 + 2
  - 以降同様
- 5個目のアイテム乗数：5個目ごとに乗数が付きます。
  - 5個目のアイテム：×2
  - 10個目のアイテム：×3
  - 以降同様

- ボーナス計算後に乗数を適用してください。

# --hints--

`getCleanupScore(["bottle", "straw", "shoe", "battery"])` は `44` を返すべきです。

```js
assert.equal(getCleanupScore(["bottle", "straw", "shoe", "battery"]), 44);
```

`getCleanupScore(["electronics", "straw", "newspaper", "bottle", "bag"])` は `58` を返すべきです。

```js
assert.equal(getCleanupScore(["electronics", "straw", "newspaper", "bottle", "bag"]), 58);
```

`getCleanupScore(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"])` は `79` を返すべきです。

```js
assert.equal(getCleanupScore(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"]), 79);
```

`getCleanupScore(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]])` は `358` を返すべきです。

```js
assert.equal(getCleanupScore(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]]), 358);
```

`getCleanupScore(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"])` は `383` を返すべきです。

```js
assert.equal(getCleanupScore(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"]), 383);
```

# --seed--

## --seed-contents--

```js
function getCleanupScore(items) {

  return items;
}
```

# --solutions--

```js
function getCleanupScore(items) {
  const values = {
    bottle: 10,
    can: 6,
    bag: 8,
    tire: 35,
    straw: 4,
    cardboard: 3,
    newspaper: 3,
    shoe: 12,
    electronics: 25,
    battery: 18,
    mattress: 38
  };

  let total = 0;
  let prevItem = null;
  let streak = 0;

  for (let i = 0; i < items.length; i++) {
    let item = items[i];
    let base;
    let isRare = Array.isArray(item) && item[0] === "rare";

    if (isRare) {
      base = item[1];
      prevItem = null;
      streak = 0;
    } else {
      base = values[item];
      streak = item === prevItem ? streak + 1 : 0;
      prevItem = item;
    }

    let points = base;
    if (!isRare) points += streak;

    if ((i + 1) % 5 === 0) {
      const multiplier = (i + 1) / 5 + 1;
      points *= multiplier;
    }

    total += points;
  }

  return total;
}
```
