---
id: 69c6ff713a52713463aa7928
title: "チャレンジ 262: URL クエリパーサー"
challengeType: 28
dashedName: challenge-262
---

# --description--

クエリ文字列を含むURLが与えられたとき、そのクエリ文字列をキーと値のペアのオブジェクト（または辞書）に解析してください。

- クエリ文字列は `"?"` の後に始まります。
- 各パラメータは `"&"` で区切られています。
- 各キーと値のペアは `"="` で区切られています。

例えば、`"https://example.com/search?name=Alice&age=30"` が与えられたら、次のように返してください：

```json
{
  "name": "Alice",
  "age": "30"
}
```

すべての値は文字列として返してください。

# --hints--

`parseUrlQuery("https://example.com/search?name=Alice&age=30")` は `{"name": "Alice", "age": "30"}` を返すはずです。

```js
assert.deepEqual(parseUrlQuery("https://example.com/search?name=Alice&age=30"), {"name": "Alice", "age": "30"});
```

`parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python")` は `{"skill": "programming", "language": "python"}` を返すはずです。

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python"), {"skill": "programming", "language": "python"});
```

`parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2")` は `{"category": "books", "sort": "asc", "page": "2"}` を返すはずです。

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2"), {"category": "books", "sort": "asc", "page": "2"});
```

`parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now")` は `{"redirect": "freecodecamp.org/learn", "when": "now"}` を返すはずです。

```js
assert.deepEqual(parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now"), {"redirect": "freecodecamp.org/learn", "when": "now"});
```

# --seed--

## --seed-contents--

```js
function parseUrlQuery(url) {

  return url;
}
```

# --solutions--

```js
function parseUrlQuery(url) {
  const queryString = url.split('?')[1];
  return Object.fromEntries(
    queryString.split('&').map(param => param.split('='))
  );
}
```
