---
id: 69c6ff713a52713463aa7926
title: "チャレンジ 260: 単語のスコア"
challengeType: 28
dashedName: challenge-260
---

# --description--

単語が与えられたら、標準の文字価値テーブルを使ってそのスコアを返してください。

| Letter | Value |
| --- | --- |
|   A |   1 |
|   B |   2 |
| ... | ... |
|   Z |  26 |

- 大文字と小文字の文字は同じ価値を持ちます。

# --hints--

`getWordScore("hi")` は `17` を返すべきです。

```js
assert.equal(getWordScore("hi"), 17);
```

`getWordScore("hello")` は `52` を返すべきです。

```js
assert.equal(getWordScore("hello"), 52);
```

`getWordScore("hippopotamus")` は `169` を返すべきです。

```js
assert.equal(getWordScore("hippopotamus"), 169);
```

`getWordScore("freeCodeCamp")` は `94` を返すべきです。

```js
assert.equal(getWordScore("freeCodeCamp"), 94);
```

# --seed--

## --seed-contents--

```js
function getWordScore(word) {

  return word;
}
```

# --solutions--

```js
function getWordScore(word) {
  return word.toLowerCase().split('').reduce((sum, char) => {
    return sum + (char.charCodeAt(0) - 96);
  }, 0);
}
```
