---
id: 69c6ff713a52713463aa7927
title: "チャレンジ 261: 数字の単語"
challengeType: 28
dashedName: challenge-261
---

# --description--

0から99までの整数が与えられたら、その英語の単語表現を返してください。

- `0` は `"zero"` を返します。
- 1から19までの数字はそれぞれ固有の名前があります（`"one"`、`"two"`、...、`"ten"`、`"eleven"`、...、`"eighteen"`、`"nineteen"`）。
- 20から90までの10の倍数はそれぞれ固有の名前があります（`"twenty"`、`"thirty"`、...、`"eighty"`、`"ninety"`）。
- 10の倍数でない21から99までの数字は、ハイフンでつながれた2語で書きます。例えば `"forty-two"` と `"fifty-three"` のように。

# --hints--

`getNumberWords(0)` は `"zero"` を返すべきです。

```js
assert.equal(getNumberWords(0), "zero");
```

`getNumberWords(10)` は `"ten"` を返すべきです。

```js
assert.equal(getNumberWords(10), "ten");
```

`getNumberWords(19)` は `"nineteen"` を返すべきです。

```js
assert.equal(getNumberWords(19), "nineteen");
```

`getNumberWords(30)` は `"thirty"` を返すべきです。

```js
assert.equal(getNumberWords(30), "thirty");
```

`getNumberWords(53)` は `"fifty-three"` を返すべきです。

```js
assert.equal(getNumberWords(53), "fifty-three");
```

`getNumberWords(7)` は `"seven"` を返すべきです。

```js
assert.equal(getNumberWords(7), "seven");
```

`getNumberWords(12)` は `"twelve"` を返すべきです。

```js
assert.equal(getNumberWords(12), "twelve");
```

`getNumberWords(60)` は `"sixty"` を返すべきです。

```js
assert.equal(getNumberWords(60), "sixty");
```

`getNumberWords(67)` は `"sixty-seven"` を返すべきです。

```js
assert.equal(getNumberWords(67), "sixty-seven");
```

`getNumberWords(98)` は `"ninety-eight"` を返すべきです。

```js
assert.equal(getNumberWords(98), "ninety-eight");
```

# --seed--

## --seed-contents--

```js
function getNumberWords(n) {

  return n;
}
```

# --solutions--

```js
function getNumberWords(n) {
  const ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
                 "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
                 "seventeen","eighteen","nineteen"];
  const tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"];

  if (n < 20) return ones[n];
  if (n % 10 === 0) return tens[Math.floor(n / 10)];
  return tens[Math.floor(n / 10)] + "-" + ones[n % 10];
}
```
