---
id: 69c6ff713a52713463aa7929
title: "チャレンジ 263: 2進数クロスワード"
challengeType: 28
dashedName: challenge-263
---

# --description--

文字が与えられたとき、その8ビットの2進数表現が以下のグリッドの中に、左右または上下のいずれかの方向で見つかるかどうかを判定してください。

```js
0 1 0 0 0 0 0 1
0 1 1 0 1 1 1 1
0 1 0 0 0 1 0 0
0 1 1 0 0 1 0 1
0 1 0 1 0 0 1 0
0 1 0 1 0 1 0 0
0 1 1 0 1 0 0 0
1 0 1 0 1 1 1 0
```

例えば、`"A"` は2進数表現 `01000001` を持ち、これは1行目の左から右にかけて現れます。

# --hints--

`isInCrossword("I")` は `true` を返すべきです。

```js
assert.isTrue(isInCrossword("I"));
```

`isInCrossword("D")` は `true` を返すべきです。

```js
assert.isTrue(isInCrossword("D"));
```

`isInCrossword("0")` は `true` を返すべきです。

```js
assert.isTrue(isInCrossword("0"));
```

`isInCrossword("u")` は `true` を返すべきです。

```js
assert.isTrue(isInCrossword("u"));
```

`isInCrossword("Y")` は `false` を返すべきです。

```js
assert.isFalse(isInCrossword("Y"));
```

`isInCrossword("p")` は `false` を返すべきです。

```js
assert.isFalse(isInCrossword("p"));
```

`isInCrossword("1")` は `false` を返すべきです。

```js
assert.isFalse(isInCrossword("1"));
```

`isInCrossword("Q")` は `false` を返すべきです。

```js
assert.isFalse(isInCrossword("Q"));
```

# --seed--

## --seed-contents--

```js
function isInCrossword(char) {

  return char;
}
```

# --solutions--

```js
function isInCrossword(char) {
  const grid = [
    "01000001",
    "01101111",
    "01000100",
    "01100101",
    "01010010",
    "01010100",
    "01101000",
    "10101110",
  ];

  const target = char.charCodeAt(0).toString(2).padStart(8, '0');
  const cols = Array.from({length: 8}, (_, c) => grid.map(r => r[c]).join(''));

  return [...grid, ...cols].some(seq => seq.includes(target) || seq.split('').reverse().join('').includes(target));
}
```
