---
id: 69cfca90e8a0a6d4d6871c51
title: "チャレンジ 267: パーセク変換器"
challengeType: 28
dashedName: challenge-267
---

# --description--

遠い銀河系では、パーセクが時間と距離の両方を測る単位として使われています。整数のパーセク数が与えられたら、それに対応する時間または距離を返してください。

- 与えられた整数が奇数なら時間を表します。偶数なら距離を表します。

以下の変換率を使ってください：

| パーセク | 時間／距離 |
| - | - |
| 1 | 2時間 |
| 2 | 6光年 |

変換した値は整数で返してください。

# --hints--

`convertParsecs(1)` は `2` を返すべきです。

```js
assert.equal(convertParsecs(1), 2);
```

`convertParsecs(2)` は `6` を返すべきです。

```js
assert.equal(convertParsecs(2), 6);
```

`convertParsecs(31)` は `62` を返すべきです。

```js
assert.equal(convertParsecs(31), 62);
```

`convertParsecs(88)` は `264` を返すべきです。

```js
assert.equal(convertParsecs(88), 264);
```

`convertParsecs(17)` は `34` を返すべきです。

```js
assert.equal(convertParsecs(17), 34);
```

`convertParsecs(14)` は `42` を返すべきです。

```js
assert.equal(convertParsecs(14), 42);
```

# --seed--

## --seed-contents--

```js
function convertParsecs(parsecs) {

  return parsecs;
}
```

# --solutions--

```js
function convertParsecs(parsecs) {
  if (parsecs % 2 !== 0) {
    return parsecs * 2;
  } else {
    return parsecs * 3;
  }
}
```
