---
id: 69cfca90e8a0a6d4d6871c4f
title: "チャレンジ 265: 最も深い括弧"
challengeType: 28
dashedName: challenge-265
---

# --description--

バランスの取れた括弧を含む文字列が与えられたとき、最も深くネストされた括弧の中身を返してください。

- 括弧は `()`、`[]`、`{}` の3種類のいずれかです。
- 入力には必ず1つだけ最も深いグループがあります。

例えば、`"(hello (world))"` が与えられたら、`"world"` を返してください。

# --hints--

`getDeepestBrackets("(hello (world))")` は `"world"` を返すべきです。

```js
assert.equal(getDeepestBrackets("(hello (world))"), "world");
```

`getDeepestBrackets("[outer [inner] outer]")` は `"inner"` を返すべきです。

```js
assert.equal(getDeepestBrackets("[outer [inner] outer]"), "inner");
```

`getDeepestBrackets("{a{b}c{d{e}f}g}")` は `"e"` を返すべきです。

```js
assert.equal(getDeepestBrackets("{a{b}c{d{e}f}g}"), "e");
```

`getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]")` は `"fox"` を返すべきです。

```js
assert.equal(getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]"), "fox");
```

`getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p")` は `"C"` を返すべきです。

```js
assert.equal(getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p"), "C");
```

# --seed--

## --seed-contents--

```js
function getDeepestBrackets(str) {

  return str;
}
```

# --solutions--

```js
function getDeepestBrackets(str) {
  let maxDepth = 0, depth = 0, content = '';
  let start = 0;

  for (let i = 0; i < str.length; i++) {
    if ('([{'.includes(str[i])) {
      depth++;
      if (depth > maxDepth) {
        maxDepth = depth;
        start = i + 1;
      }
    } else if (')]}'.includes(str[i])) {
      if (depth === maxDepth) {
        content = str.slice(start, i);
      }
      depth--;
    }
  }
  return content;
}
```
