---
id: 69cfca90e8a0a6d4d6871c50
title: "チャレンジ 266: 良い一日"
challengeType: 28
dashedName: challenge-266
---

# --description--

`"HH:MM"` 形式（24時間制）の時刻文字列が与えられたとき、次を返してください：

- `"Good morning"` から `05:00` までの時刻には `11:59`
- `"Good afternoon"` から `12:00` までの時刻には `17:59`
- `"Good evening"` から `18:00` までの時刻には `21:59`
- `"Good night"` から `22:00` までの時刻には `04:59`

# --hints--

`getGreeting("06:30")` は `"Good morning"` を返すべきです。

```js
assert.equal(getGreeting("06:30"), "Good morning");
```

`getGreeting("12:00")` は `"Good afternoon"` を返すべきです。

```js
assert.equal(getGreeting("12:00"), "Good afternoon");
```

`getGreeting("21:59")` は `"Good evening"` を返すべきです。

```js
assert.equal(getGreeting("21:59"), "Good evening");
```

`getGreeting("00:01")` は `"Good night"` を返すべきです。

```js
assert.equal(getGreeting("00:01"), "Good night");
```

`getGreeting("11:30")` は `"Good morning"` を返すべきです。

```js
assert.equal(getGreeting("11:30"), "Good morning");
```

# --seed--

## --seed-contents--

```js
function getGreeting(time) {

  return time;
}
```

# --solutions--

```js
function getGreeting(time) {
  const [hours, minutes] = time.split(':').map(Number);
  const total = hours * 60 + minutes;

  if (total >= 300 && total < 720) return 'Good morning';
  if (total >= 720 && total < 1080) return 'Good afternoon';
  if (total >= 1080 && total < 1320) return 'Good evening';
  return 'Good night';
}
```
