---
id: 69cfca90e8a0a6d4d6871c52
title: "チャレンジ 268: ナルシシスティック数"
challengeType: 28
dashedName: challenge-268
---

# --description--

正の整数がナルシシスティック数かどうか判定してください。

- ナルシシスティック数とは、各桁の数字を桁数のべき乗にしたものの合計が、その数自身と等しい数のことです。

例えば、153は3桁の数で、1<sup>3</sup> + 5<sup>3</sup> + 3<sup>3</sup> = 153となるため、ナルシシスティック数です。

# --hints--

`isNarcissistic(153)` は `true` を返すべきです。

```js
assert.isTrue(isNarcissistic(153));
```

`isNarcissistic(154)` は `false` を返すべきです。

```js
assert.isFalse(isNarcissistic(154));
```

`isNarcissistic(371)` は `true` を返すべきです。

```js
assert.isTrue(isNarcissistic(371));
```

`isNarcissistic(512)` は `false` を返すべきです。

```js
assert.isFalse(isNarcissistic(512));
```

`isNarcissistic(9)` は `true` を返すべきです。

```js
assert.isTrue(isNarcissistic(9));
```

`isNarcissistic(11)` は `false` を返すべきです。

```js
assert.isFalse(isNarcissistic(11));
```

`isNarcissistic(9474)` は `true` を返すべきです。

```js
assert.isTrue(isNarcissistic(9474));
```

`isNarcissistic(6549)` は `false` を返すべきです。

```js
assert.isFalse(isNarcissistic(6549));
```

# --seed--

## --seed-contents--

```js
function isNarcissistic(n) {

  return n;
}
```

# --solutions--

```js
function isNarcissistic(n) {
  const digits = String(n).split('');
  const power = digits.length;
  const sum = digits.reduce((acc, d) => acc + Math.pow(Number(d), power), 0);
  return sum === n;
}
```
