---
id: 69cfca90e8a0a6d4d6871c54
title: "チャレンジ 270: 最長共通部分文字列"
challengeType: 28
dashedName: challenge-270
---

# --description--

文字列が与えられたとき、複数回現れる最長の部分文字列を返してください。

- 部分文字列は重なっていてもかまいません。

# --hints--

`getLongestSubstring("abracadabra")` は `"abra"` を返すべきです。

```js
assert.equal(getLongestSubstring("abracadabra"), "abra");
```

`getLongestSubstring("hello world hello")` は `"hello"` を返すべきです。

```js
assert.equal(getLongestSubstring("hello world hello"), "hello");
```

`getLongestSubstring("mississippi")` は `"issi"` を返すべきです。

```js
assert.equal(getLongestSubstring("mississippi"), "issi");
```

`getLongestSubstring("ha ha ha ha ha ha ha")` は `"ha ha ha ha ha ha"` を返すべきです。

```js
assert.equal(getLongestSubstring("ha ha ha ha ha ha ha"), "ha ha ha ha ha ha");
```

`getLongestSubstring("the quick brown fox jumped over the lazy dog that the quick brown fox jumped over")` は `"the quick brown fox jumped over"` を返すべきです。

```js
assert.equal(getLongestSubstring("the quick brown fox jumped over the lazy dog that the quick brown fox jumped over"), "the quick brown fox jumped over");
```

# --seed--

## --seed-contents--

```js
function getLongestSubstring(str) {

  return str;
}
```

# --solutions--

```js
function getLongestSubstring(str) {
  let longest = '';

  for (let len = str.length - 1; len >= 1; len--) {
    for (let i = 0; i <= str.length - len; i++) {
      const sub = str.slice(i, i + len);
      if (str.indexOf(sub) !== str.lastIndexOf(sub)) {
        return sub;
      }
    }
  }

  return longest;
}
```
