---
id: 69d03e549613fbdbe21d11e0
title: "チャレンジ 271: 薬のリマインダー"
challengeType: 28
dashedName: challenge-271
---

# --description--

薬の配列と現在時刻を表す文字列が与えられたとき、次に服用すべき薬と服用までの時間を返してください。

- 各薬は `[name, lastTaken]` の形式で、`name` が薬の名前、`lastTaken` が最後に服用した時間です。
- すべての時間は `"HH:MM"`（24時間）形式で与えられます。

以下の薬のスケジュールを使用してください：

| 名前 | スケジュール |
| - | - |
| Deployxitrin | 08:00、16:00 |
| Debuggamanizole | 07:00、13:00、21:00 |
| Mergeflictamine | 4時間ごと |

`"{name} in Hh Mm"` の形式で文字列を返してください。例えば、`"Debuggamanizole in 2h 0m"` や `"Deployxitrin in 1h 5m"` のように。

# --hints--

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "10:00"]], "11:00")` は `"Debuggamanizole in 2h 0m"` を返すべきです。

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "10:00"]], "11:00"), "Debuggamanizole in 2h 0m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "14:55")` は `"Deployxitrin in 1h 5m"` を返すべきです。

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "14:55"), "Deployxitrin in 1h 5m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "17:15")` は `"Mergeflictamine in 0h 45m"` を返すべきです。

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "17:15"), "Mergeflictamine in 0h 45m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "09:00"]], "12:59")` は `"Debuggamanizole in 0h 1m"` を返すべきです。

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "09:00"]], "12:59"), "Debuggamanizole in 0h 1m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "21:00"], ["Mergeflictamine", "03:00"]], "06:55")` は `"Debuggamanizole in 0h 5m"` を返すべきです。

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "21:00"], ["Mergeflictamine", "03:00"]], "06:55"), "Debuggamanizole in 0h 5m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "07:30"]], "08:00")` は `"Mergeflictamine in 3h 30m"` を返すべきです。

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "07:30"]], "08:00"), "Mergeflictamine in 3h 30m");
```

# --seed--

## --seed-contents--

```js
function medicationReminder(medications, currentTime) {

  return medications;
}
```

# --solutions--

```js
function medicationReminder(medications, currentTime) {
  const schedule = {
    Deployxitrin: { type: 'fixed', times: ['08:00', '16:00'] },
    Debuggamanizole: { type: 'fixed', times: ['07:00', '13:00', '21:00'] },
    Mergeflictamine: { type: 'interval', hours: 4 }
  };

  const toMinutes = t => {
    const [h, m] = t.split(':').map(Number);
    return h * 60 + m;
  };

  const current = toMinutes(currentTime);
  let earliest = Infinity;
  let nextMed = '';

  for (const [name, lastTaken] of medications) {
    const med = schedule[name];
    let nextDue;

    if (med.type === 'fixed') {
      const futureTimes = med.times
        .map(toMinutes)
        .filter(t => t > current);
      if (futureTimes.length === 0) continue;
      nextDue = Math.min(...futureTimes);
    } else {
      nextDue = toMinutes(lastTaken) + med.hours * 60;
    }

    const diff = nextDue - current;
    if (diff > 0 && diff < earliest) {
      earliest = diff;
      nextMed = name;
    }
  }

  const h = Math.floor(earliest / 60);
  const m = earliest % 60;
  return `${nextMed} in ${h}h ${m}m`;
}
```
