---
id: 69e2383af7832c8032603b8e
title: "チャレンジ 272: 転置行列"
challengeType: 28
dashedName: challenge-272
---

# --description--

行列（配列の配列）が与えられたとき、その転置行列を返してください。

行列を転置するには、行と列を入れ替えます。例えば、インデックス `[0, 1]` にある値はインデックス `[1, 0]` に移動します。

例えば、次のような行列が与えられた場合：

```json
[
  [1, 2, 3],
  [4, 5, 6]
]
```

返すべきは：

```json
[
  [1, 4],
  [2, 5],
  [3, 6]
]
```

# --hints--

`transpose([[1, 2, 3], [4, 5, 6]])` は `[[1, 4], [2, 5], [3, 6]]` を返すべきです。

```js
assert.deepEqual(transpose([[1, 2, 3], [4, 5, 6]]), [[1, 4], [2, 5], [3, 6]]);
```

`transpose([[1, 2], [3, 4], [5, 6]])` は `[[1, 3, 5], [2, 4, 6]]` を返すべきです。

```js
assert.deepEqual(transpose([[1, 2], [3, 4], [5, 6]]), [[1, 3, 5], [2, 4, 6]]);
```

`transpose([[1, 2], [3, 4], [5, 6], [7, 8]])` は `[[1, 3, 5, 7], [2, 4, 6, 8]]` を返すべきです。

```js
assert.deepEqual(transpose([[1, 2], [3, 4], [5, 6], [7, 8]]), [[1, 3, 5, 7], [2, 4, 6, 8]]);
```

`transpose([["a", "b", "c"], ["d", "e", "f"], ["g", "h", "i"], ["j", "k", "l"]])` は `[["a", "d", "g", "j"], ["b", "e", "h", "k"], ["c", "f", "i", "l"]]` を返すべきです。

```js
assert.deepEqual(transpose([["a", "b", "c"], ["d", "e", "f"], ["g", "h", "i"], ["j", "k", "l"]]), [["a", "d", "g", "j"], ["b", "e", "h", "k"], ["c", "f", "i", "l"]]);
```

`transpose([[true, false, true, false], [false, true, false, true], [true, true, false, false], [false, false, true, true], [true, false, false, true]])` は `[[true, false, true, false, true], [false, true, true, false, false], [true, false, false, true, false], [false, true, false, true, true]]` を返すべきです。

```js
assert.deepEqual(transpose([[true, false, true, false], [false, true, false, true], [true, true, false, false], [false, false, true, true], [true, false, false, true]]), [[true, false, true, false, true], [false, true, true, false, false], [true, false, false, true, false], [false, true, false, true, true]]);
```

# --seed--

## --seed-contents--

```js
function transpose(matrix) {

  return matrix;
}
```

# --solutions--

```js
function transpose(matrix) {

  return matrix[0].map((_, i) => matrix.map(row => row[i]));
}
```
