---
id: 69e2383af7832c8032603b92
title: "チャレンジ 276: 不正な要素"
challengeType: 28
dashedName: challenge-276
---

# --description--

昇順にソートされた整数の配列があり、その中に一つだけ順序が乱れている要素があります。その要素のインデックスを返してください。

- 複数の要素が順序が乱れていると考えられる場合は、最初の要素のインデックスを返してください。

# --hints--

`findOffender([1, 6, 2, 3, 4, 5])` は `1` を返すべきです。

```js
assert.equal(findOffender([1, 6, 2, 3, 4, 5]), 1);
```

`findOffender([1, 2, 3, 5, 4, 5])` は `3` を返すべきです。

```js
assert.equal(findOffender([1, 2, 3, 5, 4, 5]), 3);
```

`findOffender([2, 1])` は `0` を返すべきです。

```js
assert.equal(findOffender([2, 1]), 0);
```

`findOffender([2, 4, 1, 6, 8])` は `2` を返すべきです。

```js
assert.equal(findOffender([2, 4, 1, 6, 8]), 2);
```

`findOffender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91])` は `6` を返すべきです。

```js
assert.equal(findOffender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91]), 6);
```

# --seed--

## --seed-contents--

```js
function findOffender(arr) {

  return arr;
}
```

# --solutions--

```js
function findOffender(arr) {
  for (let i = 0; i < arr.length; i++) {
    const without = [...arr.slice(0, i), ...arr.slice(i + 1)];
    if (without.every((n, j) => j === 0 || without[j - 1] <= n)) return i;
  }
}
```
