---
id: 69e2383af7832c8032603b94
title: "チャレンジ 278: コーヒー注文パーサー"
challengeType: 28
dashedName: challenge-278
---

# --description--

コーヒーの注文を表す文字列が与えられたとき、メニューのアイテムを特定し、フォーマットされた注文を返してください。

以下のメニューアイテムと価格を使用します：

| アイテム | 価格 | 
| - | - |
| `"cold brew"` | $4.50 |
| `"oat latte"` | $5.00 |
| `"cappuccino"` | $4.75 |
| `"espresso"` | $3.00 |
| `"vanilla syrup"` | $0.75 |
| `"caramel drizzle"` | $0.60 |
| `"extra shot"` | $0.50 |
| `"oat milk"` | $0.75 |
| `"cream"` | $0.75 |

一致したアイテムを `" + "` でつなぎ、その後にコロンとスペース（`": "`）を付けて、合計金額を返してください。

例えば、`"I'd like an oat latte with vanilla syrup and an extra shot please."` が与えられたら、`"oat latte + vanilla syrup + extra shot: $6.25"` を返します。

アイテムはメニューに出てくる順番で表示し、合計金額は常に小数点以下2桁で表示してください。

# --hints--

`formatCoffeeOrder("I'd like an oat latte with vanilla syrup and an extra shot please.")` は `"oat latte + vanilla syrup + extra shot: $6.25"` を返すべきです。

```js
assert.equal(formatCoffeeOrder("I'd like an oat latte with vanilla syrup and an extra shot please."), "oat latte + vanilla syrup + extra shot: $6.25");
```

`formatCoffeeOrder("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk.")` は `"cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85"` を返すべきです。

```js
assert.equal(formatCoffeeOrder("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk."), "cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85");
```

`formatCoffeeOrder("Can I get a cold brew with some cream and an extra shot.")` は `"cold brew + extra shot + cream: $5.75"` を返すべきです。

```js
assert.equal(formatCoffeeOrder("Can I get a cold brew with some cream and an extra shot."), "cold brew + extra shot + cream: $5.75");
```

`formatCoffeeOrder("Just an espresso please.")` は `"espresso: $3.00"` を返すべきです。

```js
assert.equal(formatCoffeeOrder("Just an espresso please."), "espresso: $3.00");
```

`formatCoffeeOrder("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle.")` は `"oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60"` を返すべきです。

```js
assert.equal(formatCoffeeOrder("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle."), "oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60");
```

# --seed--

## --seed-contents--

```js
function formatCoffeeOrder(order) {

  return order;
}
```

# --solutions--

```js
function formatCoffeeOrder(order) {
  const menu = [
    { name: "cold brew", price: 4.50 },
    { name: "oat latte", price: 5.00 },
    { name: "cappuccino", price: 4.75 },
    { name: "espresso", price: 3.00 },
    { name: "vanilla syrup", price: 0.75 },
    { name: "caramel drizzle", price: 0.60 },
    { name: "extra shot", price: 0.50 },
    { name: "oat milk", price: 0.75 },
    { name: "cream", price: 0.75 },
  ];

  const found = menu.filter(item => order.toLowerCase().includes(item.name));
  const total = found.reduce((sum, item) => sum + item.price, 0);
  return `${found.map(item => item.name).join(" + ")}: $${total.toFixed(2)}`;
}
```
