---
id: 69e2383af7832c8032603b8f
title: "チャレンジ 273: ISBN-13 バリデーター"
challengeType: 28
dashedName: challenge-273
---

# --description--

文字列が有効なISBN-13番号かどうかを判定してください。

有効なISBN-13は以下の条件を満たします：

- 数字とハイフンのみを含む
- ハイフンを除いた後にちょうど13桁の数字がある
- 次のチェックを通過する：
  1. 各数字を1または3で交互に乗算する（最初の数字は1、2番目は3、3番目は1、以下同様）。
  2. 乗算した結果の合計が10で割り切れること。

# --hints--

`isValidIsbn13("9780306406157")` は `true` を返すべきです。

```js
assert.isTrue(isValidIsbn13("9780306406157"));
```

`isValidIsbn13("97803064061570")` は `false` を返すべきです。

```js
assert.isFalse(isValidIsbn13("97803064061570"));
```

`isValidIsbn13("978-0-13-595705-9")` は `true` を返すべきです。

```js
assert.isTrue(isValidIsbn13("978-0-13-595705-9"));
```

`isValidIsbn13("978-030-64061A-4")` は `false` を返すべきです。

```js
assert.isFalse(isValidIsbn13("978-030-64061A-4"));
```

`isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9")` は `true` を返すべきです。

```js
assert.isTrue(isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn13(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn13(str) {
  const digits = str.replace(/-/g, '');
  if (!/^\d{13}$/.test(digits)) return false;
  const sum = digits.split('').reduce((acc, d, i) => acc + d * (i % 2 === 0 ? 1 : 3), 0);
  return sum % 10 === 0;
}
```
