---
id: 69f35a5bb823ed620fcb7cbc
title: "チャレンジ 283: 文字列ジッパー"
challengeType: 28
dashedName: challenge-283
---

# --description--

2つの文字列が与えられたとき、それぞれの文字を交互に組み合わせた新しい文字列を返してください。どちらかの文字列が長い場合は、残りの文字を末尾に追加してください。

最初の文字列の最初の文字から始めます。

# --hints--

`zipStrings("abc", "123")` は `"a1b2c3"` を返すべきです。

```js
assert.equal(zipStrings("abc", "123"), "a1b2c3");
```

`zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz")` は `"abcdefghijklmnopqrstuvwxyz"` を返すべきです。

```js
assert.equal(zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz"), "abcdefghijklmnopqrstuvwxyz");
```

`zipStrings("day", "night")` は `"dnaiyght"` を返すべきです。

```js
assert.equal(zipStrings("day", "night"), "dnaiyght");
```

`zipStrings("python", "javascript")` は `"pjyatvhaosncript"` を返すべきです。

```js
assert.equal(zipStrings("python", "javascript"), "pjyatvhaosncript");
```

`zipStrings("feCdCm", "reoeap")` は `"freeCodeCamp"` を返すべきです。

```js
assert.equal(zipStrings("feCdCm", "reoeap"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function zipStrings(a, b) {

  return a;
}
```

# --solutions--

```js
function zipStrings(a, b) {
  let result = '';
  const len = Math.max(a.length, b.length);
  for (let i = 0; i < len; i++) {
    if (i < a.length) result += a[i];
    if (i < b.length) result += b[i];
  }
  return result;
}
```
