---
id: 69f35a5bb823ed620fcb7cbd
title: "チャレンジ 284: I Before E のルール"
challengeType: 28
dashedName: challenge-284
---

# --description--

単語や文が与えられたとき、すべての単語が「Cの後以外はIがEの前に来る」というルールに従うように修正したバージョンを返してください。

- `"ei"` が `"c"` に続かない単語があれば、`"ie"` に置き換えてください。
- `"ie"` が `"c"` に続く単語があれば、`"ei"` に置き換えてください。
- その他の単語は変更しません。

# --hints--

`iBeforeE("beleive")` は `"believe"` を返すべきです。

```js
assert.equal(iBeforeE("beleive"), "believe");
```

`iBeforeE("recieve")` は `"receive"` を返すべきです。

```js
assert.equal(iBeforeE("recieve"), "receive");
```

`iBeforeE("we recieved a breif")` は `"we received a brief"` を返すべきです。

```js
assert.equal(iBeforeE("we recieved a breif"), "we received a brief");
```

`iBeforeE("she beleived the friendly niece could percieve the greif")` は `"she believed the friendly niece could perceive the grief"` を返すべきです。

```js
assert.equal(iBeforeE("she beleived the friendly niece could percieve the greif"), "she believed the friendly niece could perceive the grief");
```

`iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit")` は `"we received relief after the thief gave us a brief piece of fierce deceit"` を返すべきです。

```js
assert.equal(iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit"), "we received relief after the thief gave us a brief piece of fierce deceit");
```

# --seed--

## --seed-contents--

```js
function iBeforeE(sentence) {

  return sentence;
}
```

# --solutions--

```js
function iBeforeE(sentence) {
  return sentence.split(' ').map(word => {
    word = word.replace(/([^c])ei/g, '$1ie');
    word = word.replace(/cie/g, 'cei');
    return word;
  }).join(' ');
}
```
