---
id: 69f35a5bb823ed620fcb7cb8
title: "チャレンジ279：最長ドミノチェーン"
challengeType: 28
dashedName: challenge-279
---

# --description--

ドミノのセットを表す2次元配列が与えられたとき、最も長い有効なチェーンを返してください。

- 各ドミノは0～6の数字のペアで表されます。例：`[3, 2]`。
- チェーンは、各ドミノの2番目の数字が次のドミノの1番目の数字と一致するときに有効です。
- 最初のドミノの1番目の数字と最後のドミノの2番目の数字は、何とも一致する必要はありません。
- どのドミノも裏返すことができるので、`[3, 2]`は`[2, 3]`として使えます。
- 最長の有効なチェーンは必ず1つだけ存在します。

例えば、`[[1, 2], [4, 5], [2, 3]]`が与えられたら、`[[1, 2], [2, 3]]`を返してください。

# --hints--

`getLongestChain([[1, 2], [4, 5], [2, 3]])`は`[[1, 2], [2, 3]]`を返すべきです。

```js
const chain = JSON.stringify(getLongestChain([[1, 2], [4, 5], [2, 3]]));
const result1 = JSON.stringify([[1, 2], [2, 3]]);
const result2 = JSON.stringify([[3, 2], [2, 1]]);
assert.isTrue(chain === result1 || chain === result2);
```

`getLongestChain([[2, 1], [4, 3], [5, 3]])`は`[[4, 3], [3, 5]]`を返すべきです。

```js
const chain = JSON.stringify(getLongestChain([[2, 1], [4, 3], [5, 3]]));
const result1 = JSON.stringify([[4, 3], [3, 5]]);
const result2 = JSON.stringify([[5, 3], [3, 4]]);
assert.isTrue(chain === result1 || chain === result2);
```

`getLongestChain([[1, 2], [3, 4], [2, 3], [4, 0]])`は`[[1, 2], [2, 3], [3, 4], [4, 0]]`を返すべきです。

```js
const chain = JSON.stringify(getLongestChain([[1, 2], [3, 4], [2, 3], [4, 0]]));
const result1 = JSON.stringify([[1, 2], [2, 3], [3, 4], [4, 0]]);
const result2 = JSON.stringify([[0, 4], [4, 3], [3, 2], [2, 1]]);
assert.isTrue(chain === result1 || chain === result2);
```

`getLongestChain([[6, 6], [6, 1], [1, 1], [0, 3], [2, 3], [4, 1], [5, 6]])`は`[[4, 1], [1, 1], [1, 6], [6, 6], [6, 5]]`を返すべきです。

```js
const chain = JSON.stringify(getLongestChain([[6, 6], [6, 1], [1, 1], [0, 3], [2, 3], [4, 1], [5, 6]]));
const result1 = JSON.stringify([[4, 1], [1, 1], [1, 6], [6, 6], [6, 5]]);
const result2 = JSON.stringify([[5, 6], [6, 6], [6, 1], [1, 1], [1, 4]]);
assert.isTrue(chain === result1 || chain === result2);
```

`getLongestChain([[0, 4], [3, 3], [0, 3], [5, 6], [4, 5], [4, 2], [5, 5], [1, 2], [4, 4]])`は`[[3, 3], [3, 0], [0, 4], [4, 4], [4, 5], [5, 5], [5, 6]]`を返すべきです。

```js
const chain = JSON.stringify(getLongestChain([[0, 4], [3, 3], [0, 3], [5, 6], [4, 5], [4, 2], [5, 5], [1, 2], [4, 4]]));
const result1 = JSON.stringify([[3, 3], [3, 0], [0, 4], [4, 4], [4, 5], [5, 5], [5, 6]]);
const result2 = JSON.stringify([[6, 5], [5, 5], [5, 4], [4, 4], [4, 0], [0, 3], [3, 3]]);
assert.isTrue(chain === result1 || chain === result2);
```

# --seed--

## --seed-contents--

```js
function getLongestChain(dominoes) {

  return dominoes;
}
```

# --solutions--

```js
function getLongestChain(dominoes) {
  function search(chain, remaining) {
    let best = chain;
    const last = chain[chain.length - 1][1];

    for (let i = 0; i < remaining.length; i++) {
      const [a, b] = remaining[i];
      const rest = remaining.filter((_, j) => j !== i);

      if (a === last) {
        const result = search([...chain, [a, b]], rest);
        if (result.length > best.length) best = result;
      }
      if (b === last && a !== b) {
        const result = search([...chain, [b, a]], rest);
        if (result.length > best.length) best = result;
      }
    }
    return best;
  }

  let best = [];
  for (let i = 0; i < dominoes.length; i++) {
    const [a, b] = dominoes[i];
    const rest = dominoes.filter((_, j) => j !== i);

    const r1 = search([[a, b]], rest);
    if (r1.length > best.length) best = r1;

    if (a !== b) {
      const r2 = search([[b, a]], rest);
      if (r2.length > best.length) best = r2;
    }
  }
  return best;
}
```
