---
id: 69f35a5bb823ed620fcb7cb9
title: "チャレンジ 280: Mongo ID の日付"
challengeType: 28
dashedName: challenge-280
---

# --description--

MongoDBのID文字列が与えられたとき、その作成時間をISO 8601形式の文字列で返してください。

- MongoDBのIDは24文字の16進数文字列です。最初の8文字はUnixタイムスタンプ（秒単位）を16進数で表しています。

例えば、`"6a094b50bcf86cd799439011"` は16進数で`"6a094b50"` のタイムスタンプを持ち、10進数では`1778994000` で、作成時間は`"2026-05-17T05:00:00.000Z"` を表します。

# --hints--

`mongoIdToDate("6a094b50bcf86cd799439011")` は`"2026-05-17T05:00:00.000Z"` を返すべきです。

```js
assert.equal(mongoIdToDate("6a094b50bcf86cd799439011"), "2026-05-17T05:00:00.000Z");
```

`mongoIdToDate("695344eb1f4a4c1123042128")` は`"2025-12-30T03:20:11.000Z"` を返すべきです。

```js
assert.equal(mongoIdToDate("695344eb1f4a4c1123042128"), "2025-12-30T03:20:11.000Z");
```

`mongoIdToDate("386da62df34123ac54617e56")` は`"2000-01-01T07:01:01.000Z"` を返すべきです。

```js
assert.equal(mongoIdToDate("386da62df34123ac54617e56"), "2000-01-01T07:01:01.000Z");
```

`mongoIdToDate("69f571c3d7711807afd3dd55")` は`"2026-05-02T03:38:43.000Z"` を返すべきです。

```js
assert.equal(mongoIdToDate("69f571c3d7711807afd3dd55"), "2026-05-02T03:38:43.000Z");
```

`mongoIdToDate("68adce01c0e1144d0a90295a")` は`"2025-08-26T15:08:49.000Z"` を返すべきです。

```js
assert.equal(mongoIdToDate("68adce01c0e1144d0a90295a"), "2025-08-26T15:08:49.000Z");
```

# --seed--

## --seed-contents--

```js
function mongoIdToDate(id) {

  return id;
}
```

# --solutions--

```js
function mongoIdToDate(id) {
  const timestamp = parseInt(id.slice(0, 8), 16);
  return new Date(timestamp * 1000).toISOString();
}
```
