---
id: 69f35a5bb823ed620fcb7cba
title: "チャレンジ 281: ビンゴの範囲"
challengeType: 28
dashedName: challenge-281
---

# --description--

ビンゴの文字が与えられたら、その文字に対応する数字の範囲を返してください。

| 文字 | 数字の範囲 |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

最小から最大までの範囲内のすべての数字を含む配列を返してください。

# --hints--

`getBingoRange("B")` は `[1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15]` を返すべきです。

```js
assert.deepEqual(getBingoRange("B"), [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15]);
```

`getBingoRange("I")` は `[16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30]` を返すべきです。

```js
assert.deepEqual(getBingoRange("I"), [16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30]);
```

`getBingoRange("N")` は `[31, 32, 33, 34, 35, 36, 37, 38, 39, 40, 41, 42, 43, 44, 45]` を返すべきです。

```js
assert.deepEqual(getBingoRange("N"), [31, 32, 33, 34, 35, 36, 37, 38, 39, 40, 41, 42, 43, 44, 45]);
```

`getBingoRange("G")` は `[46, 47, 48, 49, 50, 51, 52, 53, 54, 55, 56, 57, 58, 59, 60]` を返すべきです。

```js
assert.deepEqual(getBingoRange("G"), [46, 47, 48, 49, 50, 51, 52, 53, 54, 55, 56, 57, 58, 59, 60]);
```

`getBingoRange("O")` は `[61, 62, 63, 64, 65, 66, 67, 68, 69, 70, 71, 72, 73, 74, 75]` を返すべきです。

```js
assert.deepEqual(getBingoRange("O"), [61, 62, 63, 64, 65, 66, 67, 68, 69, 70, 71, 72, 73, 74, 75]);
```

# --seed--

## --seed-contents--

```js
function getBingoRange(letter) {

  return letter;
}
```

# --solutions--

```js
function getBingoRange(letter) {
  const ranges = { B: [1, 15], I: [16, 30], N: [31, 45], G: [46, 60], O: [61, 75] };
  const [start, end] = ranges[letter];
  return Array.from({ length: end - start + 1 }, (_, i) => start + i);
}
```
