---
id: 69f35a5bb823ed620fcb7cbb
title: "チャレンジ 282: 睡眠負債"
challengeType: 28
dashedName: challenge-282
---

# --description--

今日までの各夜に眠った時間の配列と、1晩あたりの目標時間が与えられたとき、睡眠負債を解消するために今夜何時間眠る必要があるかを返してください。

- 今夜の睡眠時間も、追いつくために必要な合計時間に含めてください。
- 今夜の目標時間以上にすでに眠っている場合は、`0`を返してください。

# --hints--

`sleepDebt([6, 6, 6, 6, 6, 6], 8)`は`20`を返すべきです。

```js
assert.equal(sleepDebt([6, 6, 6, 6, 6, 6], 8), 20);
```

`sleepDebt([6, 7, 8, 4, 8, 6], 7)`は`10`を返すべきです。

```js
assert.equal(sleepDebt([6, 7, 8, 4, 8, 6], 7), 10);
```

`sleepDebt([10, 10, 9, 10, 9, 11], 9)`は`4`を返すべきです。

```js
assert.equal(sleepDebt([10, 10, 9, 10, 9, 11], 9), 4);
```

`sleepDebt([8, 7, 6, 7, 6, 8], 6)`は`0`を返すべきです。

```js
assert.equal(sleepDebt([8, 7, 6, 7, 6, 8], 6), 0);
```

`sleepDebt([8, 9, 10, 9, 10, 7], 7)`は`0`を返すべきです。

```js
assert.equal(sleepDebt([8, 9, 10, 9, 10, 7], 7), 0);
```

# --seed--

## --seed-contents--

```js
function sleepDebt(hoursSlept, targetHours) {

  return hoursSlept;
}
```

# --solutions--

```js
function sleepDebt(hoursSlept, targetHours) {
  const debt = (hoursSlept.length + 1) * targetHours - hoursSlept.reduce((sum, h) => sum + h, 0);
  return Math.max(0, debt);
}
```
