---
id: 69f8c998d78ad3171a0713bb
title: "チャレンジ 288: 秘密の数字"
challengeType: 28
dashedName: challenge-288
---

# --description--

秘密の数字と予想が与えられたとき、その予想が正しいかどうかを判定してください。

返す値は以下の通りです：

- 秘密の数字が予想より大きい場合は `"higher"` を返してください。
- 秘密の数字が予想より小さい場合は `"lower"` を返してください。
- 予想が正しい場合は `"you got it!"` を返してください。

# --hints--

`guessNumber(50, 30)` は `"higher"` を返すべきです。

```js
assert.equal(guessNumber(50, 30), "higher");
```

`guessNumber(85, 99)` は `"lower"` を返すべきです。

```js
assert.equal(guessNumber(85, 99), "lower");
```

`guessNumber(2026, 2026)` は `"you got it!"` を返すべきです。

```js
assert.equal(guessNumber(2026, 2026), "you got it!");
```

`guessNumber(92904, 11283)` は `"higher"` を返すべきです。

```js
assert.equal(guessNumber(92904, 11283), "higher");
```

`guessNumber(230495, 423920)` は `"lower"` を返すべきです。

```js
assert.equal(guessNumber(230495, 423920), "lower");
```

`guessNumber(120349, 120349)` は `"you got it!"` を返すべきです。

```js
assert.equal(guessNumber(120349, 120349), "you got it!");
```

# --seed--

## --seed-contents--

```js
function guessNumber(secret, guess) {

  return secret;
}
```

# --solutions--

```js
function guessNumber(secret, guess) {
  if (guess === secret) return "you got it!";
  return guess < secret ? "higher" : "lower";
}
```
