---
id: 69f8c998d78ad3171a0713b9
title: "チャレンジ 286: 未解決の課題"
challengeType: 28
dashedName: challenge-286
---

# --description--

課題番号の配列とプルリクエスト（PR）番号の配列が与えられたとき、すべてのPRがマージされた後に未解決のまま残る課題の配列を返してください。

- PRは、数字が回転した関係にある場合に課題を閉じます。たとえば、課題123はPR231や312によって閉じられます。
- PRは、まったく同じ番号の課題を閉じません。たとえば、PR123は課題123を閉じません。つまり、すべて同じ数字の課題は閉じられません。
- どちらの番号も先頭のゼロが取り除かれている場合があります。たとえば、PR201は課題12（012、201の回転）を閉じます。同様に、課題201はPR12によって閉じられます。

未解決の課題は、与えられた順序で返してください。

# --hints--

`getOpenIssues([123, 234], [231])` は `[234]` を返すべきです。

```js
assert.deepEqual(getOpenIssues([123, 234], [231]), [234]);
```

`getOpenIssues([123, 345, 16], [345, 231])` は `[345, 16]` を返すべきです。

```js
assert.deepEqual(getOpenIssues([123, 345, 16], [345, 231]), [345, 16]);
```

`getOpenIssues([456, 332, 12, 15], [201, 945, 180])` は `[456, 332, 15]` を返すべきです。

```js
assert.deepEqual(getOpenIssues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15]);
```

`getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` は `[115, 296, 24]` を返すべきです。

```js
assert.deepEqual(getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24]);
```

`getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` は `[95, 395, 754, 296, 709, 237, 1802]` を返すべきです。

```js
assert.deepEqual(getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802]);
```

# --seed--

## --seed-contents--

```js
function getOpenIssues(issues, prs) {

  return issues;
}
```

# --solutions--

```js
function getOpenIssues(issues, prs) {
  function getRotations(n) {
    const s = String(n);
    const rotations = new Set();
    for (let i = 0; i < s.length; i++) {
      const rotated = s.slice(i) + s.slice(0, i);
      rotations.add(Number(rotated));
    }
    return rotations;
  }

  function areRotations(a, b) {
    if (a === b) return false;
    const rotations = getRotations(a);
    return rotations.has(b);
  }

  return issues.filter(issue =>
    !prs.some(pr => areRotations(pr, issue) || areRotations(issue, pr))
  );
}
```
