---
id: 69f8c998d78ad3171a0713bf
title: "チャレンジ 292: より広いアスペクト比"
challengeType: 28
dashedName: challenge-292
---

# --description--

異なる画像の寸法を表す2つの文字列が与えられたとき、幅と高さの比率がより大きい画像のアスペクト比を返してください。

- 与えられる文字列は `"WxH"` の形式で、例えば `"1920x1080"` のようになります。
- アスペクト比は幅と高さの比率で、最も簡単な整数比に約分します。例えば `"1920x1080"` は `"16:9"` に約分されます。
- `"W:H"` の形式の文字列を返してください。例えば `"16:9"` のようにします。

# --hints--

`getWiderAspectRatio("1920x1080", "800x600")` は `"16:9"` を返すべきです。

```js
assert.equal(getWiderAspectRatio("1920x1080", "800x600"), "16:9");
```

`getWiderAspectRatio("1080x1350", "2048x1536")` は `"4:3"` を返すべきです。

```js
assert.equal(getWiderAspectRatio("1080x1350", "2048x1536"), "4:3");
```

`getWiderAspectRatio("640x480", "2440x1220")` は `"2:1"` を返すべきです。

```js
assert.equal(getWiderAspectRatio("640x480", "2440x1220"), "2:1");
```

`getWiderAspectRatio("360x640", "1080x1920")` は `"9:16"` を返すべきです。

```js
assert.equal(getWiderAspectRatio("360x640", "1080x1920"), "9:16");
```

`getWiderAspectRatio("3440x1440", "2048x858")` は `"43:18"` を返すべきです。

```js
assert.equal(getWiderAspectRatio("3440x1440", "2048x858"), "43:18");
```

`getWiderAspectRatio("12345x61234", "12534x51234")` は `"2089:8539"` を返すべきです。

```js
assert.equal(getWiderAspectRatio("12345x61234", "12534x51234"), "2089:8539");
```

# --seed--

## --seed-contents--

```js
function getWiderAspectRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function getWiderAspectRatio(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  function parse(str) {
    const [w, h] = str.split("x").map(Number);
    const g = gcd(w, h);
    return { ratio: w / h, simplified: `${w / g}:${h / g}` };
  }

  const imageA = parse(a);
  const imageB = parse(b);

  return imageA.ratio >= imageB.ratio ? imageA.simplified : imageB.simplified;
}
```
