---
id: 69f8c998d78ad3171a0713ba
title: "チャレンジ 287: ローマ数字の修正"
challengeType: 28
dashedName: challenge-287
---

# --description--

不正なローマ数字の文字列が与えられたとき、標準的なローマ数字表記でその値を返してください。

入力は加算表記のみを使用するため、各記号はその値を合計に加えます。参考までに、記号と値は以下の通りです：

| 記号 | 値 |
| ------ | ----- |
| `"I"`  | 1     |
| `"V"`  | 5     |
| `"X"`  | 10    |
| `"L"`  | 50    |
| `"C"`  | 100   |
| `"D"`  | 500   |
| `"M"`  | 1000  |

再エンコードする際は、各ステップで可能な限り大きな記号を使い、必要に応じて減算表記の組み合わせ（`"IV"`、`"IX"`、`"XL"`、`"XC"`、`"CD"`、`"CM"`）を使用してください。

# --hints--

`fixNumerals("XIIIII")` は `"XV"` を返すべきです。

```js
assert.equal(fixNumerals("XIIIII"), "XV");
```

`fixNumerals("IIIILX")` は `"LXIV"` を返すべきです。

```js
assert.equal(fixNumerals("IIIILX"), "LXIV");
```

`fixNumerals("XXVVVIIIII")` は `"XL"` を返すべきです。

```js
assert.equal(fixNumerals("XXVVVIIIII"), "XL");
```

`fixNumerals("MDCCLXXXXVIIII")` は `"MDCCXCIX"` を返すべきです。

```js
assert.equal(fixNumerals("MDCCLXXXXVIIII"), "MDCCXCIX");
```

`fixNumerals("IIIIVVVVXXXXLLLLCCDD")` は `"MCDLXIV"` を返すべきです。

```js
assert.equal(fixNumerals("IIIIVVVVXXXXLLLLCCDD"), "MCDLXIV");
```

`fixNumerals("ILCDMIVDIIXLCVCXDL")` は `"MMCMLXXXIV"` を返すべきです。

```js
assert.equal(fixNumerals("ILCDMIVDIIXLCVCXDL"), "MMCMLXXXIV");
```

# --seed--

## --seed-contents--

```js
function fixNumerals(str) {

  return str;
}
```

# --solutions--

```js
function fixNumerals(str) {
  const values = { I: 1, V: 5, X: 10, L: 50, C: 100, D: 500, M: 1000 };
  const encoding = [
    [1000, "M"], [900, "CM"], [500, "D"], [400, "CD"],
    [100, "C"], [90, "XC"], [50, "L"], [40, "XL"],
    [10, "X"], [9, "IX"], [5, "V"], [4, "IV"], [1, "I"]
  ];

  let total = [...str].reduce((sum, c) => sum + values[c], 0);

  let result = "";
  for (const [value, symbol] of encoding) {
    while (total >= value) {
      result += symbol;
      total -= value;
    }
  }
  return result;
}
```
