---
id: 69f8c998d78ad3171a0713bc
title: "チャレンジ 289: 差の合計"
challengeType: 28
dashedName: challenge-289
---

# --description--

数値の配列が与えられたとき、各数値とその次の数値の差の合計を返してください。

例えば、`[1, 3, 4]` が与えられたら、`3`（2 + 1）を返します。

# --hints--

`sumOfDifferences([1, 3, 4])` は `3` を返すべきです。

```js
assert.equal(sumOfDifferences([1, 3, 4]), 3);
```

`sumOfDifferences([5, -3, 3, 9, 10])` は `5` を返すべきです。

```js
assert.equal(sumOfDifferences([5, -3, 3, 9, 10]), 5);
```

`sumOfDifferences([9, 6, 15, -20, 33, 14, 25, 16, -7])` は `-16` を返すべきです。

```js
assert.equal(sumOfDifferences([9, 6, 15, -20, 33, 14, 25, 16, -7]), -16);
```

`sumOfDifferences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75])` は `25` を返すべきです。

```js
assert.equal(sumOfDifferences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75]), 25);
```

# --seed--

## --seed-contents--

```js
function sumOfDifferences(arr) {

  return arr;
}
```

# --solutions--

```js
function sumOfDifferences(arr) {
  return arr.slice(0, -1).reduce((sum, n, i) => sum + (arr[i + 1] - n), 0);
}
```
