---
id: 69f8c998d78ad3171a0713bd
title: "チャレンジ 290: ピザパーティー"
challengeType: 28
dashedName: challenge-290
---

# --description--

今日の各人の労働時間が入った配列が与えられたとき、ピザパーティーのために注文するピザの枚数を返してください。

- 各人の労働時間を3で割って、その人のスライス数を求めます。
- スライスは部分的に食べられないので、各人のスライス数は切り上げて整数にしてください。
- 各人には最低2スライスが割り当てられます。
- ピザ1枚は8スライスです。合計のピザ枚数は切り上げて整数にしてください。

# --hints--

`getPizzasToOrder([8, 8, 8])` は `2` を返すべきです。

```js
assert.equal(getPizzasToOrder([8, 8, 8]), 2);
```

`getPizzasToOrder([10, 9, 8, 2, 2, 6, 10])` は `3` を返すべきです。

```js
assert.equal(getPizzasToOrder([10, 9, 8, 2, 2, 6, 10]), 3);
```

`getPizzasToOrder([1, 2, 3, 4, 5])` は `2` を返すべきです。

```js
assert.equal(getPizzasToOrder([1, 2, 3, 4, 5]), 2);
```

`getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8])` は `3` を返すべきです。

```js
assert.equal(getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8]), 3);
```

`getPizzasToOrder([9, 9, 6])` は `1` を返すべきです。

```js
assert.equal(getPizzasToOrder([9, 9, 6]), 1);
```

`getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0])` は `5` を返すべきです。

```js
assert.equal(getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5);
```

# --seed--

## --seed-contents--

```js
function getPizzasToOrder(hoursWorked) {

  return hoursWorked;
}
```

# --solutions--

```js
function getPizzasToOrder(hoursWorked) {
  const totalSlices = hoursWorked.reduce((sum, hours) => {
    return sum + Math.max(Math.ceil(hours / 3), 2);
  }, 0);

  return Math.ceil(totalSlices / 8);
}
```
