---
id: 69f90c1329a94b37e2a2086d
title: "チャレンジ294：丸括弧の組み合わせ"
challengeType: 28
dashedName: challenge-294
---

# --description--

整数`n`が与えられたとき、`n`組の丸括弧の有効な組み合わせの数を返してください。

- 有効な組み合わせとは、すべての開き丸括弧に対応する閉じ丸括弧があり、対応する開き丸括弧より前に閉じ丸括弧が現れない文字列のことです。

例えば、`2`が与えられた場合、有効な組み合わせは2つあります。

```json
(())
()()
```

# --hints--

`getCombinations(2)`は`2`を返す必要があります。

```js
assert.equal(getCombinations(2), 2);
```

`getCombinations(3)`は`5`を返す必要があります。

```js
assert.equal(getCombinations(3), 5);
```

`getCombinations(5)`は`42`を返す必要があります。

```js
assert.equal(getCombinations(5), 42);
```

`getCombinations(8)`は`1430`を返す必要があります。

```js
assert.equal(getCombinations(8), 1430);
```

`getCombinations(13)`は`742900`を返す必要があります。

```js
assert.equal(getCombinations(13), 742900);
```

# --seed--

## --seed-contents--

```js
function getCombinations(n) {

  return n;
}
```

# --solutions--

```js
function getCombinations(n) {
  let count = 0;

  function generate(open, close) {
    if (open === 0 && close === 0) {
      count++;
      return;
    }
    if (open > 0) generate(open - 1, close);
    if (close > open) generate(open, close - 1);
  }

  generate(n, n);
  return count;
}
```
