---
id: 69f90c1329a94b37e2a2086c
title: "チャレンジ 293: ベストハンド"
challengeType: 28
dashedName: challenge-293
---

# --description--

5つの文字列の配列で表されたトランプのカードから、最も強い役の名前を返してください。

- 各カードは2文字の文字列で表されます：ランクの後にスートが続きます。`"2h"` 例えば。
  - ランクは低い順に、`"2"`、`"3"`、`"4"`、`"5"`、`"6"`、`"7"`、`"8"`、`"9"`、`"T"`、`"J"`、`"Q"`、`"K"`、`"A"` です。
  - スートは `"h"`、`"d"`、`"c"`、`"s"` です。
- エース（`"A"`）はストレートの中で高いランクまたは低いランクとして使うことができます。

役は、弱い順から強い順に以下の通りです：

| 名前 | 説明 |
| - | - |
| `"High Card"` | ペアなしまたはそれ以下 |
| `"Pair"` | 1つのランクの2枚 |
| `"Two Pair"` | 2つの異なるランクの2枚ずつ |
| `"Three of a Kind"` | 1つのランクの3枚 |
| `"Straight"` | 連続した5つのランク |
| `"Flush"` | 同じスートの5枚 |
| `"Full House"` | 1つのランクの3枚と別のランクの2枚 |
| `"Four of a Kind"` | 1つのランクの4枚 |
| `"Straight Flush"` | 同じスートで連続した5つのランク |
| `"Royal Flush"` | 同じスートの `"A"`、`"K"`、`"Q"`、`"J"`、`"T"` |

最も強い役の名前を返してください。

# --hints--

`getBestHand(["7s", "7h", "7d", "2c", "5h"])` は `"Three of a Kind"` を返すべきです。

```js
assert.equal(getBestHand(["7s", "7h", "7d", "2c", "5h"]), "Three of a Kind");
```

`getBestHand(["Ks", "Kh", "Kd", "4s", "4h"])` は `"Full House"` を返すべきです。

```js
assert.equal(getBestHand(["Ks", "Kh", "Kd", "4s", "4h"]), "Full House");
```

`getBestHand(["2h", "5h", "7h", "9h", "Jh"])` は `"Flush"` を返すべきです。

```js
assert.equal(getBestHand(["2h", "5h", "7h", "9h", "Jh"]), "Flush");
```

`getBestHand(["As", "Ah", "Ad", "Ac", "Kh"])` は `"Four of a Kind"` を返すべきです。

```js
assert.equal(getBestHand(["As", "Ah", "Ad", "Ac", "Kh"]), "Four of a Kind");
```

`getBestHand(["Ts", "Th", "9d", "9c", "8h"])` は `"Two Pair"` を返すべきです。

```js
assert.equal(getBestHand(["Ts", "Th", "9d", "9c", "8h"]), "Two Pair");
```

`getBestHand(["9c", "8c", "7c", "6c", "5c"])` は `"Straight Flush"` を返すべきです。

```js
assert.equal(getBestHand(["9c", "8c", "7c", "6c", "5c"]), "Straight Flush");
```

`getBestHand(["As", "Kh", "Jd", "8c", "5h"])` は `"High Card"` を返すべきです。

```js
assert.equal(getBestHand(["As", "Kh", "Jd", "8c", "5h"]), "High Card");
```

`getBestHand(["As", "2h", "3d", "4c", "5h"])` は `"Straight"` を返すべきです。

```js
assert.equal(getBestHand(["As", "2h", "3d", "4c", "5h"]), "Straight");
```

`getBestHand(["Ts", "Th", "7c", "6d", "5h"])` は `"Pair"` を返すべきです。

```js
assert.equal(getBestHand(["Ts", "Th", "7c", "6d", "5h"]), "Pair");
```

`getBestHand(["As", "Ks", "Qs", "Js", "Ts"])` は `"Royal Flush"` を返すべきです。

```js
assert.equal(getBestHand(["As", "Ks", "Qs", "Js", "Ts"]), "Royal Flush");
```

# --seed--

## --seed-contents--

```js
function getBestHand(cards) {

  return cards;
}
```

# --solutions--

```js
function getBestHand(cards) {
  const rankOrder = "23456789TJQKA";

  const ranks = cards.map(c => c[0]);
  const suits = cards.map(c => c[1]);

  const isFlush = suits.every(s => s === suits[0]);

  const rankIndices = ranks.map(r => rankOrder.indexOf(r)).sort((a, b) => a - b);

  function isStraight(indices) {
    const isRegular = indices.every((r, i) => i === 0 || r === indices[i - 1] + 1);
    const isLowAce = JSON.stringify(indices) === JSON.stringify([0, 1, 2, 3, 12]);
    return isRegular || isLowAce;
  }

  const straight = isStraight(rankIndices);

  const counts = {};
  ranks.forEach(r => counts[r] = (counts[r] || 0) + 1);
  const freq = Object.values(counts).sort((a, b) => b - a);

  if (isFlush && JSON.stringify(rankIndices) === JSON.stringify([8, 9, 10, 11, 12])) return "Royal Flush";
  if (isFlush && straight) return "Straight Flush";
  if (freq[0] === 4) return "Four of a Kind";
  if (freq[0] === 3 && freq[1] === 2) return "Full House";
  if (isFlush) return "Flush";
  if (straight) return "Straight";
  if (freq[0] === 3) return "Three of a Kind";
  if (freq[0] === 2 && freq[1] === 2) return "Two Pair";
  if (freq[0] === 2) return "Pair";
  return "High Card";
}
```
