---
id: 6a0dcd03ee4e68698080ef6c
title: "チャレンジ 308: クレジットカードのバリデーター"
challengeType: 28
dashedName: challenge-308
---

# --description--

クレジットカード番号の数字の文字列が与えられたとき、次の方法を使って有効なカード番号かどうかを判定してください。

- 後ろから2番目の数字から始めて、左に向かって1つおきに数字を2倍にします。
- 2倍にした数字が9より大きい場合は、9を引きます。
- すべての数字（2倍にしたものとそのままのもの）を合計します。
- 合計が10で割り切れれば、その番号は有効です。

# --hints--

`isValidCard("4532015112830366")` は `true` を返すべきです。

```js
assert.isTrue(isValidCard("4532015112830366"));
```

`isValidCard("5425233430109903")` は `true` を返すべきです。

```js
assert.isTrue(isValidCard("5425233430109903"));
```

`isValidCard("371449635398431")` は `true` を返すべきです。

```js
assert.isTrue(isValidCard("371449635398431"));
```

`isValidCard("6011111111111117")` は `true` を返すべきです。

```js
assert.isTrue(isValidCard("6011111111111117"));
```

`isValidCard("4532015112830367")` は `false` を返すべきです。

```js
assert.isFalse(isValidCard("4532015112830367"));
```

`isValidCard("1234567890123456")` は `false` を返すべきです。

```js
assert.isFalse(isValidCard("1234567890123456"));
```

`isValidCard("4532015112830368")` は `false` を返すべきです。

```js
assert.isFalse(isValidCard("4532015112830368"));
```

# --seed--

## --seed-contents--

```js
function isValidCard(number) {

  return number;
}
```

# --solutions--

```js
function isValidCard(number) {
  let sum = 0;

  for (let i = number.length - 1; i >= 0; i--) {
    let digit = parseInt(number[i]);

    if ((number.length - 1 - i) % 2 === 1) {
      digit *= 2;
      if (digit > 9) digit -= 9;
    }

    sum += digit;
  }

  return sum % 10 === 0;
}
```
