---
id: 6a0dcc730cb92a616f86f0c0
title: "チャレンジ 296: スキーマバリデーター パート2"
challengeType: 28
dashedName: challenge-296
---

# --description--

オブジェクト（JavaScriptの場合）または辞書（Pythonの場合）が、以下のスキーマに一致するかどうかを判定してください。

```json
{
  username: string,
  posts: number,
  verified: boolean
}
```

- 追加のキーは許可されます

# --hints--

`isValidSchema({ username: "alice", posts: 10, verified: false })` は `true` を返すべきです。

```js
assert.isTrue(isValidSchema({ username: "alice", posts: 10, verified: false }));
```

`isValidSchema({ username: "carol", posts: 15, verified: true, followers: 25 })` は `true` を返すべきです。

```js
assert.isTrue(isValidSchema({ username: "carol", posts: 15, verified: true, followers: 25 }));
```

`isValidSchema({ username: "frank", posts: "21", verified: true })` は `false` を返すべきです。

```js
assert.isFalse(isValidSchema({ username: "frank", posts: "21", verified: true }));
```

`isValidSchema({ username: "sam", posts: 17, verified: "false" })` は `false` を返すべきです。

```js
assert.isFalse(isValidSchema({ username: "sam", posts: 17, verified: "false" }));
```

`isValidSchema({ username: "bill", verified: true })` は `false` を返すべきです。

```js
assert.isFalse(isValidSchema({ username: "bill", verified: true }));
```

`isValidSchema({ username: "fred", verified: true })` は `false` を返すべきです。

```js
assert.isFalse(isValidSchema({ username: "fred", verified: true }));
```

`isValidSchema({ username: 5, posts: 10, verified: true })` は `false` を返すべきです。

```js
assert.isFalse(isValidSchema({ username: 5, posts: 10, verified: true }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  return (
    typeof obj.username === 'string' &&
    typeof obj.posts === 'number' &&
    typeof obj.verified === 'boolean'
  );
}
```
