---
id: 6a0dcc730cb92a616f86f0bf
title: "チャレンジ 295: スキーマバリデーター パート1"
challengeType: 28
dashedName: challenge-295
---

# --description--

オブジェクト（JavaScript）または辞書（Python）が、以下のスキーマに合致するかどうかを判定してください。

```json
{
  username: string
}
```

- 追加のキーは許可されます

# --hints--

`isValidSchema({ username: "bob" })` は `true` を返すべきです。

```js
assert.isTrue(isValidSchema({ username: "bob" }));
```

`isValidSchema({ username: "jen", posts: 30 })` は `true` を返すべきです。

```js
assert.isTrue(isValidSchema({ username: "jen", posts: 30 }));
```

`isValidSchema({ username: "" })` は `true` を返すべきです。

```js
assert.isTrue(isValidSchema({ username: "" }));
```

`isValidSchema({ username: 7 })` は `false` を返すべきです。

```js
assert.isFalse(isValidSchema({ username: 7 }));
```

`isValidSchema({ posts: 25 })` は `false` を返すべきです。

```js
assert.isFalse(isValidSchema({ posts: 25 }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  return typeof obj.username === 'string';
}
```
