---
id: 6a0dcd03ee4e68698080ef68
title: "チャレンジ 304: 旅程の手配"
challengeType: 28
dashedName: challenge-304
---

# --description--

1日旅行のための少なくとも2つのオプションのストップが含まれる配列が与えられたとき、有効な旅程の手配数を返してください。

旅程には常に `"breakfast"`、`"lunch"`、および `"dinner"` が含まれ、これらは引数として渡されません。オプションのストップは旅程のどこにでも配置できますが、以下のルールに従う必要があります。

- `"breakfast"` は常に最初で、`"lunch"` の前に少なくとも1つのストップがあります。
- `"lunch"` は `"dinner"` の前に現れ、間に少なくとも1つのストップがあります。
- `"dinner"` の後には最大で1つのオプションのストップしか現れません。

有効な手配の数を返してください。

# --hints--

`getItineraryCount(["library", "park"])` は `2` を返すべきです。

```js
assert.equal(getItineraryCount(["library", "park"]), 2);
```

`getItineraryCount(["library", "park", "arcade"])` は `18` を返すべきです。

```js
assert.equal(getItineraryCount(["library", "park", "arcade"]), 18);
```

`getItineraryCount(["library", "park", "arcade", "store"])` は `120` を返すべきです。

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store"]), 120);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe"])` は `840` を返すべきです。

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe"]), 840);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"])` は `55440` を返すべきです。

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440);
```

# --seed--

## --seed-contents--

```js
function getItineraryCount(stops) {

  return stops;
}
```

# --solutions--

```js
function getItineraryCount(stops) {
  const items = ["lunch", "dinner", ...stops];
  let count = 0;

  function permute(arr) {
    if (arr.length === 0) return [[]];
    return arr.flatMap((item, i) =>
      permute([...arr.slice(0, i), ...arr.slice(i + 1)]).map(p => [item, ...p])
    );
  }

  for (const perm of permute(items)) {
    const lunchIdx = perm.indexOf("lunch");
    const dinnerIdx = perm.indexOf("dinner");

    if (lunchIdx >= dinnerIdx) continue;
    if (lunchIdx < 1) continue;
    if (dinnerIdx - lunchIdx < 2) continue;
    if (perm.length - dinnerIdx - 1 > 1) continue;

    count++;
  }

  return count;
}
```
