---
id: 6a0dcc730cb92a616f86f0c5
title: "チャレンジ 301: 最後の残り"
challengeType: 28
dashedName: challenge-301
---

# --description--

残っている洗濯用洗剤のスクープ数と、過去の日ごとに使ったスクープ数の配列が与えられたとき、残っている洗剤で何日分の洗濯ができるかを返してください。

使用履歴から1日あたりの平均使用量を計算し、今後もその量を毎日使うと仮定してください。

# --hints--

`lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2])` は `5` を返すべきです。

```js
assert.equal(lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2]), 5);
```

`lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1])` は `7` を返すべきです。

```js
assert.equal(lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1]), 7);
```

`lastLoadDate(33, [5, 0, 4, 3, 3, 2])` は `11` を返すべきです。

```js
assert.equal(lastLoadDate(33, [5, 0, 4, 3, 3, 2]), 11);
```

`lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2])` は `12` を返すべきです。

```js
assert.equal(lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2]), 12);
```

`lastLoadDate(20, [13, 9, 12, 10, 8])` は `1` を返すべきです。

```js
assert.equal(lastLoadDate(20, [13, 9, 12, 10, 8]), 1);
```

# --seed--

## --seed-contents--

```js
function lastLoadDate(scoops, usage) {

  return scoops;
}
```

# --solutions--

```js
function lastLoadDate(scoops, usage) {
  const avg = usage.reduce((sum, n) => sum + n, 0) / usage.length;
  return Math.floor(scoops / avg);
}
```
