---
id: 6a0dcd03ee4e68698080ef6b
title: "チャレンジ 307: ゾーニング規制"
challengeType: 28
dashedName: challenge-307
---

# --description--

2次元のグリッド（配列の配列）で表された都市の建物配置が与えられたとき、ゾーニング規制に違反しているすべての建物の座標を返してください。

グリッドの各セルには、以下のテーブルのラベルのいずれかが含まれています。建物は、その（最大）4つの隣接セル（水平または垂直）が隣接できないタイプである場合、違反しています。

| ラベル | 種類 | 隣接できないもの |
| - | - | - |
| `"i"` | 工業 | `"R"`, `"I"` |
| `"A"` | 農業 | `"C"` |
| `"R"` | 住宅 | `"i"`, `"C"` |
| `"I"` | 公共施設 | `"i"` |
| `"C"` | 商業 | `"R"`, `"A"` |
| `""` (空文字) | 未開発 | 制限なし |

違反しているすべてのセルの座標を、任意の順序で`[row, col]`ペアの配列として返してください。違反がなければ、空の配列を返してください。

# --hints--

`getZoneViolations([["R", "C"], ["", "C"]])`は`[[0, 0], [0, 1]]`を返すべきです。

```js
assert.deepEqual(getZoneViolations([["R", "C"], ["", "C"]]).sort(), [[0, 0], [0, 1]]);
```

`getZoneViolations([["", "i"], ["", "R"], ["R", "I"]])`は`[[0, 1], [1, 1]]`を返すべきです。

```js
assert.deepEqual(getZoneViolations([["", "i"], ["", "R"], ["R", "I"]]).sort(), [[0, 1], [1, 1]]);
```

`getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])`は`[]`を返すべきです。

```js
assert.deepEqual(getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]]).sort(), []);
```

`getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])`は`[[0, 1], [0, 2], [0, 3]]`を返すべきです。

```js
assert.deepEqual(getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]]).sort(), [[0, 1], [0, 2], [0, 3]]);
```


`getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])`は`[[2, 3], [2, 4], [3, 1], [3, 2]]`を返すべきです。

```js
assert.deepEqual(getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]]).sort(), [[2, 3], [2, 4], [3, 1], [3, 2]]);
```

# --seed--

## --seed-contents--

```js
function getZoneViolations(grid) {

  return grid;
}
```

# --solutions--

```js
function getZoneViolations(grid) {
  const rules = {
    "i": ["R", "I"],
    "A": ["C"],
    "R": ["i", "C"],
    "I": ["i"],
    "C": ["R", "A"],
    "": []
  };

  const directions = [[-1, 0], [0, -1], [0, 1], [1, 0]];
  const violations = [];

  for (let row = 0; row < grid.length; row++) {
    for (let col = 0; col < grid[row].length; col++) {
      const cell = grid[row][col];
      const forbidden = rules[cell];

      for (const [dr, dc] of directions) {
        const nr = row + dr;
        const nc = col + dc;
        if (nr >= 0 && nr < grid.length && nc >= 0 && nc < grid[row].length) {
          if (forbidden.includes(grid[nr][nc])) {
            violations.push([row, col]);
            break;
          }
        }
      }
    }
  }

  return violations;
}
```
