---
id: 6a0dcd03ee4e68698080ef6d
title: "チャレンジ 309: 数字の並べ替え"
challengeType: 28
dashedName: challenge-309
---

# --description--

カンマで区切られた数字の文字列が与えられたら、数字を小さい順に並べ替えた配列を返してください。

# --hints--

`sortNumbers("3,1,2")` は `[1, 2, 3]` を返す必要があります。

```js
assert.deepEqual(sortNumbers("3,1,2"), [1, 2, 3]);
```

`sortNumbers("5,3,8,1,9,2")` は `[1, 2, 3, 5, 8, 9]` を返す必要があります。

```js
assert.deepEqual(sortNumbers("5,3,8,1,9,2"), [1, 2, 3, 5, 8, 9]);
```

`sortNumbers("12,61,49,80,19,50,77,38")` は `[12, 19, 38, 49, 50, 61, 77, 80]` を返す必要があります。

```js
assert.deepEqual(sortNumbers("12,61,49,80,19,50,77,38"), [12, 19, 38, 49, 50, 61, 77, 80]);
```

`sortNumbers("0,6,-19,44,-2,7,0")` は `[-19, -2, 0, 0, 6, 7, 44]` を返す必要があります。

```js
assert.deepEqual(sortNumbers("0,6,-19,44,-2,7,0"), [-19, -2, 0, 0, 6, 7, 44]);
```

# --seed--

## --seed-contents--

```js
function sortNumbers(str) {

  return str;
}
```

# --solutions--

```js
function sortNumbers(str) {
  return str.split(",").map(Number).sort((a, b) => a - b);
}
```
