---
id: 6a0dcd03ee4e68698080ef66
title: "チャレンジ 302: ジェットラグ"
challengeType: 28
dashedName: challenge-302
---

# --description--

出発都市、到着都市、飛行時間（時間単位）、移動方向が与えられたとき、旅行者が経験しているジェットラグの時間数を返してください。

与えられる都市は以下のリストの中からで、それぞれのUTCオフセットが含まれています：

| 都市 | オフセット |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

ジェットラグ時間の計算方法：

1. 2つの都市間のタイムゾーン差（時間単位）を求めます。
2. 移動方向の乗数を決定します。`"east"`に移動する場合は1.5、それ以外は1.0です。
3. ジェットラグ時間は次の式で求めます：タイムゾーン差 + （飛行時間 × 0.1）× 移動方向の乗数

ジェットラグ時間は小数点以下1桁に丸めて返してください。

# --hints--

`getJetLagHours("Istanbul", "Hong Kong", 10, "east")` は `6.5` を返すべきです。

```js
assert.equal(getJetLagHours("Istanbul", "Hong Kong", 10, "east"), 6.5);
```

`getJetLagHours("London", "New York", 8, "west")` は `5.8` を返すべきです。

```js
assert.equal(getJetLagHours("London", "New York", 8, "west"), 5.8);
```

`getJetLagHours("Hong Kong", "Tokyo", 4, "east")` は `1.6` を返すべきです。

```js
assert.equal(getJetLagHours("Hong Kong", "Tokyo", 4, "east"), 1.6);
```

`getJetLagHours("Dubai", "London", 7, "west")` は `4.7` を返すべきです。

```js
assert.equal(getJetLagHours("Dubai", "London", 7, "west"), 4.7);
```

`getJetLagHours("Los Angeles", "Hong Kong", 15, "west")` は `17.5` を返すべきです。

```js
assert.equal(getJetLagHours("Los Angeles", "Hong Kong", 15, "west"), 17.5);
```

`getJetLagHours("Tokyo", "Dubai", 9, "west")` は `5.9` を返すべきです。

```js
assert.equal(getJetLagHours("Tokyo", "Dubai", 9, "west"), 5.9);
```

`getJetLagHours("New York", "Istanbul", 10, "east")` は `9.5` を返すべきです。

```js
assert.equal(getJetLagHours("New York", "Istanbul", 10, "east"), 9.5);
```

# --seed--

## --seed-contents--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {

  return departureCity;
}
```

# --solutions--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {
  const offsets = {
    "Los Angeles": -8,
    "New York": -5,
    "London": 0,
    "Istanbul": +3,
    "Dubai": +4,
    "Hong Kong": +8,
    "Tokyo": +9
  };

  const depOffset = offsets[departureCity];
  const arrOffset = offsets[arrivalCity];

  const timezoneDiff = Math.abs(arrOffset - depOffset);
  const multiplier = direction === "east" ? 1.5 : 1.0;

  const score = timezoneDiff + (flightDuration * 0.1) * multiplier;
  return Math.round(score * 10) / 10;
}
```
