---
id: 6a15cadf5f240d05a264955b
title: "チャレンジ 317: BMI計算機"
challengeType: 28
dashedName: challenge-317
---

# --description--

ポンドでの体重とインチでの身長が与えられたら、小数点以下1桁に丸めたBMI（ボディマス指数）を返してください。

BMIを求めるには、体重を身長の二乗で割り、その結果に703を乗算してください。

# --hints--

`calculateBmi(180, 70)` は `25.8` を返すべきです。

```js
assert.equal(calculateBmi(180, 70), 25.8);
```

`calculateBmi(140, 64)` は `24.0` を返すべきです。

```js
assert.equal(calculateBmi(140, 64), 24.0);
```

`calculateBmi(160, 76)` は `19.5` を返すべきです。

```js
assert.equal(calculateBmi(160, 76), 19.5);
```

`calculateBmi(200, 60)` は `39.1` を返すべきです。

```js
assert.equal(calculateBmi(200, 60), 39.1);
```

`calculateBmi(150, 68)` は `22.8` を返すべきです。

```js
assert.equal(calculateBmi(150, 68), 22.8);
```

# --seed--

## --seed-contents--

```js
function calculateBmi(weight, height) {

  return weight;
}
```

# --solutions--

```js
function calculateBmi(weight, height) {
  return Math.round((weight / (height * height)) * 703 * 10) / 10;
}
```
