---
id: 6a15cadf5f240d05a2649559
title: "チャレンジ 315: 夏至"
challengeType: 28
dashedName: challenge-315
---

# --description--

今日は夏至で、北半球では一年で最も昼が長く、南半球では最も短い日です。緯度が与えられたら、昼間と夜間の時間を表す文字列を返してください。

- 緯度は90（北極）から-90（南極）までの範囲です
- 昼間の時間数 = 12 + (緯度 / 90) * 12
- 昼間の時間数は最も近い偶数に丸めてください

昼間の時間は`"☀️"`、夜間の時間は`"🌑"`を使った24文字の文字列を返してください。条件は以下の通りです：

- 各文字は1時間を表し、午前0時（0時）から始まります
- 日の出と日の入りは正午を中心に対称になります

例えば、緯度`0`（赤道）は昼間が12時間なので、日の出は午前6時、日の入りは午後6時です。返す文字列は`"🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑"`です。

# --hints--

`getDaytimeHours(0)`は`"🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑"`を返すべきです。

```js
assert.equal(getDaytimeHours(0), "🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑");
```

`getDaytimeHours(90)`は`"☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️"`を返すべきです。

```js
assert.equal(getDaytimeHours(90), "☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️");
```

`getDaytimeHours(-90)`は`"🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑"`を返すべきです。

```js
assert.equal(getDaytimeHours(-90), "🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑");
```

`getDaytimeHours(-33)`は`"🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑"`を返すべきです。

```js
assert.equal(getDaytimeHours(-33), "🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑");
```

`getDaytimeHours(66.5)`は`"🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑"`を返すべきです。

```js
assert.equal(getDaytimeHours(66.5), "🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑");
```

`getDaytimeHours(40)`は`"🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑"`を返すべきです。

```js
assert.equal(getDaytimeHours(40), "🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑");
```

`getDaytimeHours(-50)`は`"🌑🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑🌑"`を返すべきです。

```js
assert.equal(getDaytimeHours(-50), "🌑🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑🌑");
```

# --seed--

## --seed-contents--

```js
function getDaytimeHours(latitude) {

  return latitude;
}
```

# --solutions--

```js
function getDaytimeHours(latitude) {
  const daylightHours = Math.round((12 + (latitude / 90) * 12) / 2) * 2;
  const sunrise = 12 - daylightHours / 2;
  const sunset = 12 + daylightHours / 2;

  let result = "";
  for (let hour = 0; hour < 24; hour++) {
    result += hour >= sunrise && hour < sunset ? "☀️" : "🌑";
  }

  return result;
}
```
