---
id: 6a15cadf5f240d05a264955e
title: "チャレンジ 320: 血液バンク"
challengeType: 28
dashedName: challenge-320
---

# --description--

血液バンクの在庫の配列と患者の血液型リクエストの配列が与えられたとき、`"X of Y patients served"` の形式の文字列を返してください。ここで `X` はバンクの在庫から血液を受け取れる最大の患者数、`Y` は患者の総数です。

両方の配列の各エントリーは、`"AB"`、`"A"`、`"B"`、または `"O"` のいずれかの血液型です。

適合ルール：

- `"AB"` はどの血液型からも受け取れます。
- `"A"` は `"A"` と `"O"` から受け取れます。
- `"B"` は `"B"` と `"O"` から受け取れます。
- `"O"` は `"O"` からのみ受け取れます。

与えられた配列の重複エントリーは数量を表します。

# --hints--

`triageBlood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"])` は `"4 of 4 patients served"` を返すべきです。

```js
assert.equal(triageBlood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"]), "4 of 4 patients served");
```

`triageBlood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"])` は `"3 of 5 patients served"` を返すべきです。

```js
assert.equal(triageBlood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"]), "3 of 5 patients served");
```

`triageBlood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"])` は `"4 of 5 patients served"` を返すべきです。

```js
assert.equal(triageBlood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"]), "4 of 5 patients served");
```

`triageBlood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"])` は `"4 of 4 patients served"` を返すべきです。

```js
assert.equal(triageBlood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"]), "4 of 4 patients served");
```

`triageBlood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"])` は `"8 of 11 patients served"` を返すべきです。

```js
assert.equal(triageBlood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"]), "8 of 11 patients served");
```

`triageBlood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"])` は `"13 of 14 patients served"` を返すべきです。

```js
assert.equal(triageBlood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"]), "13 of 14 patients served");
```

# --seed--

## --seed-contents--

```js
function triageBlood(bank, patients) {

  return bank;
}
```

# --solutions--

```js
function triageBlood(bank, patients) {
  const b = {};
  const p = {};
  for (const t of bank) b[t] = (b[t] || 0) + 1;
  for (const t of patients) p[t] = (p[t] || 0) + 1;

  const serve = (patient, donors) => {
    for (const donor of donors) {
      const n = Math.min(p[patient] || 0, b[donor] || 0);
      p[patient] = (p[patient] || 0) - n;
      b[donor] = (b[donor] || 0) - n;
      saved += n;
    }
  };

  let saved = 0;
  serve("O",  ["O"]);
  serve("A",  ["A", "O"]);
  serve("B",  ["B", "O"]);
  serve("AB", ["AB", "A", "B", "O"]);
  return `${saved} of ${patients.length} patients served`;
}
```
