---
id: 6a15cadf5f240d05a2649557
title: "チャレンジ 313: レンタル料金"
challengeType: 28
dashedName: challenge-313
---

# --description--

レンタルのタイムスタンプ、返却のタイムスタンプ、レンタルの階層が与えられたとき、遅延料金を含むレンタルの合計料金を返してください。

- 与えられるタイムスタンプはUTCのISO文字列です。例：`"2026-06-18T18:30:00Z"`。
- レンタルの階層はレンタルの返却期限までの日数で、`1`、`3`、または`7`です。
- レンタルはレンタル期間の最終日の12:00 PM UTCまで、またはそれ以前に返却する必要があります。例えば、3月15日にいつでも借りた1日レンタルは、3月16日の12:00 PM UTCまでに返却しなければなりません。
- 返却期限を過ぎた日数ごとに遅延料金が発生します。

料金は以下の通りです：

| 階層 | 基本料金 | 遅延料金（日ごと） |
| - | - | - |
| 1日 | $4.99 | $3.99 |
| 3日 | $3.99 | $2.99 |
| 7日 | $2.99 | $0.99 |

合計料金は小数点以下2桁に丸めて、`"$D.CC"`の形式で返してください。

# --hints--

`getRentalCost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1)` は `"$4.99"` を返すべきです。

```js
assert.equal(getRentalCost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1), "$4.99");
```

`getRentalCost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1)` は `"$12.97"` を返すべきです。

```js
assert.equal(getRentalCost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1), "$12.97");
```

`getRentalCost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3)` は `"$3.99"` を返すべきです。

```js
assert.equal(getRentalCost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3), "$3.99");
```

`getRentalCost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3)` は `"$9.97"` を返すべきです。

```js
assert.equal(getRentalCost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3), "$9.97");
```

`getRentalCost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7)` は `"$2.99"` を返すべきです。

```js
assert.equal(getRentalCost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7), "$2.99");
```

`getRentalCost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7)` は `"$358.40"` を返すべきです。

```js
assert.equal(getRentalCost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7), "$358.40");
```

# --seed--

## --seed-contents--

```js
function getRentalCost(rented, returned, tier) {

  return rented;
}
```

# --solutions--

```js
function getRentalCost(rented, returned, tier) {
  const pricing = {
    1: { base: 4.99, late: 3.99 },
    3: { base: 3.99, late: 2.99 },
    7: { base: 2.99, late: 0.99 },
  };

  const rentedDate = new Date(rented);
  const returnDate = new Date(returned);
  const due = new Date(rentedDate);
  due.setUTCDate(due.getUTCDate() + tier);
  due.setUTCHours(12, 0, 0, 0);
  const { base, late } = pricing[tier];

  if (returnDate <= due) {
    return `$${base.toFixed(2)}`;
  }

  const msPerDay = 1000 * 60 * 60 * 24;
  const daysLate = Math.ceil((returnDate - due) / msPerDay);
  const total = base + daysLate * late;

  return `$${total.toFixed(2)}`;
}
```
