---
id: 6a15cadf5f240d05a2649558
title: "チャレンジ 314: 素因数分解"
challengeType: 28
dashedName: challenge-314
---

# --description--

1より大きい整数が与えられたら、その素因数分解を昇順の数値の配列として返してください。

素因数分解とは、与えられた整数を乗算して作る素数の集合です。各整数には正確に1つの集合があります。例えば、20の素因数分解は `[2, 2, 5]` です。なぜなら 2 * 2 * 5 = 20 だからです。

与えられた整数が素数そのものであれば、単一要素の配列で返してください。

# --hints--

`primeFactorization(20)` は `[2, 2, 5]` を返すべきです。

```js
assert.deepEqual(primeFactorization(20), [2, 2, 5]);
```

`primeFactorization(17)` は `[17]` を返すべきです。

```js
assert.deepEqual(primeFactorization(17), [17]);
```

`primeFactorization(15)` は `[3, 5]` を返すべきです。

```js
assert.deepEqual(primeFactorization(15), [3, 5]);
```

`primeFactorization(35)` は `[5, 7]` を返すべきです。

```js
assert.deepEqual(primeFactorization(35), [5, 7]);
```

`primeFactorization(999)` は `[3, 3, 3, 37]` を返すべきです。

```js
assert.deepEqual(primeFactorization(999), [3, 3, 3, 37]);
```

`primeFactorization(360)` は `[2, 2, 2, 3, 3, 5]` を返すべきです。

```js
assert.deepEqual(primeFactorization(360), [2, 2, 2, 3, 3, 5]);
```

`primeFactorization(510510)` は `[2, 3, 5, 7, 11, 13, 17]` を返すべきです。

```js
assert.deepEqual(primeFactorization(510510), [2, 3, 5, 7, 11, 13, 17]);
```

# --seed--

## --seed-contents--

```js
function primeFactorization(n) {

  return n;
}
```

# --solutions--

```js
function primeFactorization(n) {
  const factors = [];
  let divisor = 2;

  while (divisor * divisor <= n) {
    while (n % divisor === 0) {
      factors.push(divisor);
      n /= divisor;
    }
    divisor++;
  }

  if (n > 1) factors.push(n);

  return factors;
}
```
