---
id: 6a15cadf5f240d05a264955a
title: "チャレンジ 316: 1337 スピーク"
challengeType: 28
dashedName: challenge-316
---

# --description--

小文字の文字列が与えられたら、以下の文字をそれぞれのリート置換に変換してリートスピークに翻訳したものを返してください。

| 文字 | リート |
| - | - |
| `a` | `4` |
| `e` | `3` |
| `g` | `9` |
| `i` | `1` |
| `l` | `1` |
| `o` | `0` |
| `s` | `5` |
| `t` | `7` |

- 置換がない文字はそのままにしてください。

# --hints--

`makeLeet("cool")` は `"c001"` を返すべきです。

```js
assert.equal(makeLeet("cool"), "c001");
```

`makeLeet("leet")` は `"1337"` を返すべきです。

```js
assert.equal(makeLeet("leet"), "1337");
```

`makeLeet("hacker")` は `"h4ck3r"` を返すべきです。

```js
assert.equal(makeLeet("hacker"), "h4ck3r");
```

`makeLeet("satellite")` は `"547311173"` を返すべきです。

```js
assert.equal(makeLeet("satellite"), "547311173");
```

`makeLeet("abcdefghijklmnopqrstuvwxyz")` は `"4bcd3f9h1jk1mn0pqr57uvwxyz"` を返すべきです。

```js
assert.equal(makeLeet("abcdefghijklmnopqrstuvwxyz"), "4bcd3f9h1jk1mn0pqr57uvwxyz");
```

# --seed--

## --seed-contents--

```js
function makeLeet(str) {

  return str;
}
```

# --solutions--

```js
function makeLeet(str) {
  const map = { a: "4", e: "3", g: "9", i: "1", l: "1", o: "0", t: "7", s: "5" };
  return str.split("").map(char => map[char] || char).join("");
}
```
