---
id: 6a151d32d772271dd2448c2d
title: "チャレンジ 311: 呪文使い"
challengeType: 28
dashedName: challenge-311
---

# --description--

あなたが唱えている呪文コードの文字列が与えられたとき、合計スコアを計算してください。

文字列の各文字は呪文を表します：

| Code | 呪文 | カテゴリー | 基本スコア |
| - | - | - | - |
| `"f"` | 火 | 破壊 | 3 |
| `"l"` | 雷 | 破壊 | 3 |
| `"i"` | 氷 | 制御 | 2 |
| `"w"` | 風 | 制御 | 2 |
| `"h"` | 回復 | 回復 | 1 |
| `"s"` | 盾 | 回復 | 1 |

異なるカテゴリーの呪文が連続して唱えられた回数に応じてコンボ倍率が適用されます：

- 最初の呪文は常に基本スコアで計算されます。
- 前の呪文と異なるカテゴリーの呪文が連続するたびに倍率が1ずつ増加します。
- 前の呪文と同じカテゴリーの呪文を唱えると倍率は1にリセットされます。
- 各呪文のスコアは基本スコアに現在の倍率を掛けたものです。

呪文のシーケンスから合計スコアを返してください。

# --hints--

`cast("fihwl")` は `33` を返すべきです。

```js
assert.equal(cast("fihwl"), 33);
```

`cast("lwswfi")` は `45` を返すべきです。

```js
assert.equal(cast("lwswfi"), 45);
```

`cast("wislhfl")` は `37` を返すべきです。

```js
assert.equal(cast("wislhfl"), 37);
```

`cast("sihwlih")` は `50` を返すべきです。

```js
assert.equal(cast("sihwlih"), 50);
```

`cast("wishlfihwslwifihl")` は `101` を返すべきです。

```js
assert.equal(cast("wishlfihwslwifihl"), 101);
```

# --seed--

## --seed-contents--

```js
function cast(spells) {

  return spells;
}
```

# --solutions--

```js
function cast(spells) {
  const lookup = {
    f: { category: "destruction", score: 3 },
    l: { category: "destruction", score: 3 },
    i: { category: "control", score: 2 },
    w: { category: "control", score: 2 },
    h: { category: "restoration", score: 1 },
    s: { category: "restoration", score: 1 }
  };

  let total = 0;
  let multiplier = 1;
  let prevCategory = null;

  for (const spell of spells) {
    const { category, score } = lookup[spell];

    if (prevCategory && category !== prevCategory) {
      multiplier++;
    } else {
      multiplier = 1;
    }

    total += score * multiplier;
    prevCategory = category;
  }

  return total;
}
```
