---
id: 6a15cadf5f240d05a264955d
title: "チャレンジ 319: Frontmatter パーサー"
challengeType: 28
dashedName: challenge-319
---

# --description--

frontmatter ブロックを表す文字列が与えられたら、それを解析してキーと値を持つオブジェクト（JavaScript）または辞書（Python）を返してください。

Frontmatter は `---` で囲まれており、その中に1行ごとに `key: value` のペアが含まれています。例えば：

```md
---
title: My Post
draft: false
views: 100
---
```

次のように返す必要があります：

```js
{
  title: "My Post",
  draft: false,
  views: 100
}
```

- 数値、ブール値、文字列はそれぞれの型として返してください。
- 与えられる文字列は改行文字（`"\n"`）で区切られています。上記の例は `"---\ntitle: My Post\ndraft: false\nviews: 100\n---"` のように与えられます。

# --hints--

`parseFrontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---")` は `{ title: "My Post", draft: false, views: 100 }` を返す必要があります。

```js
assert.deepEqual(parseFrontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---"), { title: "My Post", draft: false, views: 100 });
```

`parseFrontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---")` は `{ id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: false }` を返す必要があります。

```js
assert.deepEqual(parseFrontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---"), { id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: false });
```

`parseFrontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---")` は `{ version: "1.0.0", url: "https://example.com", private: true }` を返す必要があります。

```js
assert.deepEqual(parseFrontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---"), { version: "1.0.0", url: "https://example.com", private: true });
```

`parseFrontmatter("---\nrating: 4.5\nprice: 9.99\n---")` は `{ rating: 4.5, price: 9.99 }` を返す必要があります。

```js
assert.deepEqual(parseFrontmatter("---\nrating: 4.5\nprice: 9.99\n---"), { rating: 4.5, price: 9.99 });
```

# --seed--

## --seed-contents--

```js
function parseFrontmatter(str) {

  return str;
}
```

# --solutions--

```js
function parseFrontmatter(str) {
  const lines = str.split("\n").filter(line => line !== "---" && line !== "");

  return lines.reduce((obj, line) => {
    const [key, ...rest] = line.split(":");
    const value = rest.join(":").trim();

    if (value === "true") obj[key.trim()] = true;
    else if (value === "false") obj[key.trim()] = false;
    else if (!isNaN(value)) obj[key.trim()] = Number(value);
    else obj[key.trim()] = value;

    return obj;
  }, {});
}
```
