---
id: 6a19b1062d1b153d8ac76d72
title: "チャレンジ 324: 重複文字のカウント"
challengeType: 28
dashedName: challenge-324
---

# --description--

2つの文字列が与えられたとき、2番目の文字列にある文字のうち、1番目の文字列に見つかる文字の数を返してください。

- 2番目の文字列に重複している文字は、それぞれ別々にカウントされます。

# --hints--

`duplicateCharacterCount("aloha", "hei")` は `1` を返すべきです。

```js
assert.equal(duplicateCharacterCount("aloha", "hei"), 1);
```

`duplicateCharacterCount("jambo", "bonjour")` は `4` を返すべきです。

```js
assert.equal(duplicateCharacterCount("jambo", "bonjour"), 4);
```

`duplicateCharacterCount("hello", "hola")` は `3` を返すべきです。

```js
assert.equal(duplicateCharacterCount("hello", "hola"), 3);
```

`duplicateCharacterCount("ola", "hej")` は `0` を返すべきです。

```js
assert.equal(duplicateCharacterCount("ola", "hej"), 0);
```

`duplicateCharacterCount("ciao", "konnichiwa")` は `5` を返すべきです。

```js
assert.equal(duplicateCharacterCount("ciao", "konnichiwa"), 5);
```

`duplicateCharacterCount("merhaba", "xin chao")` は `2` を返すべきです。

```js
assert.equal(duplicateCharacterCount("merhaba", "xin chao"), 2);
```

`duplicateCharacterCount("hello world", "hello to everyone around the world")` は `26` を返すべきです。

```js
assert.equal(duplicateCharacterCount("hello world", "hello to everyone around the world"), 26);
```

# --seed--

## --seed-contents--

```js
function duplicateCharacterCount(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function duplicateCharacterCount(str1, str2) {
  return str2.split("").filter(char => str1.includes(char)).length;
}
```
