---
id: 6a19dd69388e5a6e59e4f2ef
title: "チャレンジ 326: 最大利益"
challengeType: 28
dashedName: challenge-326
---

# --description--

日々の株価の配列と予算（ドル単位）が与えられたとき、その期間内に株を売買して得られる最大の利益を計算してください。

- 売るのは買った後だけにしてください。
- 買いと売りの取引は最大で1回ずつまで行えます。売った後は再び買うことはできません。
- 株は丸ごとの単位でしか買えません。
- 最大の利益を文字列で返してください。利益は1セント単位で切り捨て、小数点以下2桁の形式にしてください。

# --hints--

`getMaxProfit([5, 6], 50)` は `"10.00"` を返すべきです。

```js
assert.equal(getMaxProfit([5, 6], 50), "10.00");
```

`getMaxProfit([8, 2, 5, 10], 20)` は `"80.00"` を返すべきです。

```js
assert.equal(getMaxProfit([8, 2, 5, 10], 20), "80.00");
```

`getMaxProfit([4, 5, 3, 6], 20)` は `"18.00"` を返すべきです。

```js
assert.equal(getMaxProfit([4, 5, 3, 6], 20), "18.00");
```

`getMaxProfit([54.40, 51.22, 53.99, 50.28, 53.01, 52.84], 200)` は `"8.31"` を返すべきです。

```js
assert.equal(getMaxProfit([54.40, 51.22, 53.99, 50.28, 53.01, 52.84], 200), "8.31");
```

`getMaxProfit([15.38, 15.01, 14.99, 14.62, 14.28], 80)` は `"0.00"` を返すべきです。

```js
assert.equal(getMaxProfit([15.38, 15.01, 14.99, 14.62, 14.28], 80), "0.00");
```

`getMaxProfit([121.45, 126.82, 122.91, 124.65, 128.83, 128.83, 127.33], 1230.25)` は `"73.80"` を返すべきです。

```js
assert.equal(getMaxProfit([121.45, 126.82, 122.91, 124.65, 128.83, 128.83, 127.33], 1230.25), "73.80");
```

# --seed--

## --seed-contents--

```js
function getMaxProfit(prices, budget) {

  return prices;
}
```

# --solutions--

```js
function getMaxProfit(prices, budget) {
  let maxProfit = 0;

  for (let buy = 0; buy < prices.length - 1; buy++) {
    for (let sell = buy + 1; sell < prices.length; sell++) {
      const shares = Math.floor(budget / prices[buy]);
      const profit = shares * (prices[sell] - prices[buy]);
      if (profit > maxProfit) maxProfit = profit;
    }
  }

  return (Math.floor(maxProfit * 100) / 100).toFixed(2);
}
```
