---
id: 6a19b1062d1b153d8ac76d73
title: "チャレンジ 325: ラッキーナンバー"
challengeType: 28
dashedName: challenge-325
---

# --description--

人の名前のファーストネームとラストネームの文字列が与えられたとき、次のルールを使ってラッキーナンバーを計算してください。

- ファーストネームとラストネームはスペースで区切られています
- それぞれの名前の母音と子音の数を数えます
- 小さい方の母音と子音の数を掛け合わせ、それに小さい方の名前の長さを掛けます
- 大きい方の母音と子音の数と大きい方の名前の長さでも同じ計算をします
- 大きい方の値から小さい方の値を引いてラッキーナンバーを求めます

最終的な値がゼロ（`0`）の場合は、`13`を返してください。

# --hints--

`getLuckyNumber("John Doe")` は `21` を返すはずです。

```js
assert.equal(getLuckyNumber("John Doe"), 21);
```

`getLuckyNumber("Olivia Lewis")` は `52` を返すはずです。

```js
assert.equal(getLuckyNumber("Olivia Lewis"), 52);
```

`getLuckyNumber("James Wilson")` は `18` を返すはずです。

```js
assert.equal(getLuckyNumber("James Wilson"), 18);
```

`getLuckyNumber("Elizabeth Hernandez")` は `81` を返すはずです。

```js
assert.equal(getLuckyNumber("Elizabeth Hernandez"), 81);
```

`getLuckyNumber("Mike Walker")` は `32` を返すはずです。

```js
assert.equal(getLuckyNumber("Mike Walker"), 32);
```

`getLuckyNumber("Chloe Perez")` は `13` を返すはずです。

```js
assert.equal(getLuckyNumber("Chloe Perez"), 13);
```

# --seed--

## --seed-contents--

```js
function getLuckyNumber(name) {

  return name;
}
```

# --solutions--

```js
function getLuckyNumber(name) {
  const vowels = "aeiouAEIOU";

  function nameStats(str) {
    const vowelCount = str.split("").filter(c => vowels.includes(c)).length;
    const consonantCount = str.split("").filter(c => c.match(/[a-zA-Z]/) && !vowels.includes(c)).length;
    return { vowelCount, consonantCount, length: str.length };
  }

  const [first, last] = name.split(" ");
  const firstStats = nameStats(first);
  const lastStats = nameStats(last);

  const largerVowels = Math.max(firstStats.vowelCount, lastStats.vowelCount);
  const smallerVowels = Math.min(firstStats.vowelCount, lastStats.vowelCount);
  const largerConsonants = Math.max(firstStats.consonantCount, lastStats.consonantCount);
  const smallerConsonants = Math.min(firstStats.consonantCount, lastStats.consonantCount);
  const largerLength = Math.max(firstStats.length, lastStats.length);
  const smallerLength = Math.min(firstStats.length, lastStats.length);

  const largerValue = largerVowels * largerConsonants * largerLength;
  const smallerValue = smallerVowels * smallerConsonants * smallerLength;
  const result = largerValue - smallerValue;

  return result === 0 ? 13 : result;
}
```
