---
id: 6a1d9f98e819ed70a0e994e2
title: "チャレンジ 336: ホロスコープの相性"
challengeType: 28
dashedName: challenge-336
---

# --description--

2つの星座の文字列が与えられたら、その相性のパーセンテージを返してください。

星座は12の位置のホイールに次の順で並んでいます：`"Aries"`、`"Taurus"`、`"Gemini"`、`"Cancer"`、`"Leo"`、`"Virgo"`、`"Libra"`、`"Scorpio"`、`"Sagittarius"`、`"Capricorn"`、`"Aquarius"`、`"Pisces"`。`"Pisces"`の後は`"Aries"`に戻ります。2つの星座間の最短距離を見つけて、相性を返してください：

| 距離 | 相性 |
|----------|---------------|
| 0 | `"100%"` |
| 1 | `"40%"` |
| 2 | `"80%"` |
| 3 | `"30%"` |
| 4 | `"90%"` |
| 5 | `"20%"` |
| 6 | `"50%"` |

# --hints--

`horoscopeMatch("Libra", "Sagittarius")` は `"80%"` を返すべきです。

```js
assert.equal(horoscopeMatch("Libra", "Sagittarius"), "80%");
```

`horoscopeMatch("Gemini", "Scorpio")` は `"20%"` を返すべきです。

```js
assert.equal(horoscopeMatch("Gemini", "Scorpio"), "20%");
```

`horoscopeMatch("Pisces", "Aries")` は `"40%"` を返すべきです。

```js
assert.equal(horoscopeMatch("Pisces", "Aries"), "40%");
```

`horoscopeMatch("Capricorn", "Cancer")` は `"50%"` を返すべきです。

```js
assert.equal(horoscopeMatch("Capricorn", "Cancer"), "50%");
```

`horoscopeMatch("Aquarius", "Aquarius")` は `"100%"` を返すべきです。

```js
assert.equal(horoscopeMatch("Aquarius", "Aquarius"), "100%");
```

`horoscopeMatch("Virgo", "Taurus")` は `"90%"` を返すべきです。

```js
assert.equal(horoscopeMatch("Virgo", "Taurus"), "90%");
```

`horoscopeMatch("Leo", "Scorpio")` は `"30%"` を返すべきです。

```js
assert.equal(horoscopeMatch("Leo", "Scorpio"), "30%");
```

# --seed--

## --seed-contents--

```js
function horoscopeMatch(sign1, sign2) {

  return sign1;
}
```

# --solutions--

```js
function horoscopeMatch(sign1, sign2) {
  const signs = ["Aries", "Taurus", "Gemini", "Cancer", "Leo", "Virgo", "Libra", "Scorpio", "Sagittarius", "Capricorn", "Aquarius", "Pisces"];
  
  const a = signs.indexOf(sign1) + 1;
  const b = signs.indexOf(sign2) + 1;
  const diff = Math.abs(a - b);
  const distance = Math.min(diff, 12 - diff);

  const compatibility = { 0: "100%", 1: "40%", 2: "80%", 3: "30%", 4: "90%", 5: "20%", 6: "50%" };
  return compatibility[distance];
}
```
