---
id: 6a1d9f98e819ed70a0e994da
title: "チャレンジ 328: Kaprekarのルーチン"
challengeType: 28
dashedName: challenge-328
---

# --description--

4桁の数字が与えられたとき、6174に到達するまでにKaprekarのルーチンを適用する回数を返してください。

Kaprekarのルーチンは次のように動作します：

- 桁を降順に並べて最大の数字を作る
- 桁を昇順に並べて最小の数字を作る（必要に応じて先頭にゼロを追加してください）
- 小さい方の数字を大きい方から引く
- 新しい数字で繰り返す

# --hints--

`kaprekar(1234)` は `3` を返すべきです。

```js
assert.equal(kaprekar(1234), 3);
```

`kaprekar(2025)` は `6` を返すべきです。

```js
assert.equal(kaprekar(2025), 6);
```

`kaprekar(7173)` は `4` を返すべきです。

```js
assert.equal(kaprekar(7173), 4);
```

`kaprekar(3164)` は `7` を返すべきです。

```js
assert.equal(kaprekar(3164), 7);
```

`kaprekar(8082)` は `2` を返すべきです。

```js
assert.equal(kaprekar(8082), 2);
```

# --seed--

## --seed-contents--

```js
function kaprekar(n) {

  return n;
}
```

# --solutions--

```js
function kaprekar(n) {
  let steps = 0;
  let current = n;

  while (current !== 6174) {
    const digits = String(current).padStart(4, "0").split("").map(Number);
    const desc = parseInt(digits.slice().sort((a, b) => b - a).join(""));
    const asc = parseInt(digits.slice().sort((a, b) => a - b).join(""));
    current = desc - asc;
    steps++;
  }

  return steps;
}
```
