---
id: 6a1d9f98e819ed70a0e994dc
title: "チャレンジ 330: 小文字の単語"
challengeType: 28
dashedName: challenge-330
---

# --description--

文字列が与えられたら、完全に小文字の単語だけを元の順序で返し、各単語の間にスペースを入れてください。

# --hints--

`getLowercaseWords("hello GOOD world")` は `"hello world"` を返すべきです。

```js
assert.equal(getLowercaseWords("hello GOOD world"), "hello world");
```

`getLowercaseWords("these are all lowercase")` は `"these are all lowercase"` を返すべきです。

```js
assert.equal(getLowercaseWords("these are all lowercase"), "these are all lowercase");
```

`getLowercaseWords("less is NoT more")` は `"less is more"` を返すべきです。

```js
assert.equal(getLowercaseWords("less is NoT more"), "less is more");
```

`getLowercaseWords("DonT eat pizza every OTHER day")` は `"eat pizza every day"` を返すべきです。

```js
assert.equal(getLowercaseWords("DonT eat pizza every OTHER day"), "eat pizza every day");
```

`getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog")` は `"the quick brown fox jumped over the lazy dog"` を返すべきです。

```js
assert.equal(getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog"), "the quick brown fox jumped over the lazy dog");
```

# --seed--

## --seed-contents--

```js
function getLowercaseWords(str) {

  return str;
}
```

# --solutions--

```js
function getLowercaseWords(str) {
  return str.split(" ").filter(word => word === word.toLowerCase()).join(" ");
}
```
