---
id: 6a1d9f98e819ed70a0e994e0
title: "チャレンジ 334: 正確なおつり"
challengeType: 28
dashedName: challenge-334
---

# --description--

整数のセント単位の金額が与えられたとき、ペニー（1セント）、ニッケル（5セント）、ダイム（10セント）、クォーター（25セント）を使って正確なおつりを作る異なる方法の数を返してください。

# --hints--

`exactChange(3)` は `1` を返すべきです。

```js
assert.equal(exactChange(3), 1);
```

`exactChange(9)` は `2` を返すべきです。

```js
assert.equal(exactChange(9), 2);
```

`exactChange(17)` は `6` を返すべきです。

```js
assert.equal(exactChange(17), 6);
```

`exactChange(39)` は `24` を返すべきです。

```js
assert.equal(exactChange(39), 24);
```

`exactChange(61)` は `73` を返すべきです。

```js
assert.equal(exactChange(61), 73);
```

`exactChange(99)` は `213` を返すべきです。

```js
assert.equal(exactChange(99), 213);
```

# --seed--

## --seed-contents--

```js
function exactChange(amount) {

  return amount;
}
```

# --solutions--

```js
function exactChange(amount) {
  const coins = [1, 5, 10, 25];
  const dp = new Array(amount + 1).fill(0);
  dp[0] = 1;

  for (const coin of coins) {
    for (let i = coin; i <= amount; i++) {
      dp[i] += dp[i - coin];
    }
  }

  return dp[amount];
}
```
