---
id: 6a1d9f98e819ed70a0e994e1
title: "チャレンジ 335: 5つのサイコロ"
challengeType: 28
dashedName: challenge-335
---

# --description--

1～6の値を持つ5つのサイコロの配列が与えられたら、可能な限り最良の役を返してください。

役は最も低いものから高いものへ次のようにランク付けされています：

| 役 | 説明 |
|------|-------------|
| `"no pair"` | ペアなしまたはそれ以上の役なし |
| `"pair"` | 同じ値のサイコロが2つ |
| `"two pair"` | 2つの異なるペア |
| `"three of a kind"` | 同じ値のサイコロが3つ |
| `"small straight"` | 4つの連続した値 |
| `"large straight"` | 5つの連続した値 |
| `"full house"` | 3つの同じ値とペア |
| `"four of a kind"` | 同じ値のサイコロが4つ |
| `"five of a kind"` | 5つすべてのサイコロが同じ値 |

# --hints--

`fiveDice([1, 1, 1, 1, 1])` は `"five of a kind"` を返すべきです。

```js
assert.equal(fiveDice([1, 1, 1, 1, 1]), "five of a kind");
```

`fiveDice([5, 5, 5, 6, 5])` は `"four of a kind"` を返すべきです。

```js
assert.equal(fiveDice([5, 5, 5, 6, 5]), "four of a kind");
```

`fiveDice([2, 5, 6, 4, 3])` は `"large straight"` を返すべきです。

```js
assert.equal(fiveDice([2, 5, 6, 4, 3]), "large straight");
```

`fiveDice([4, 3, 3, 3, 1])` は `"three of a kind"` を返すべきです。

```js
assert.equal(fiveDice([4, 3, 3, 3, 1]), "three of a kind");
```

`fiveDice([4, 6, 2, 6, 5])` は `"pair"` を返すべきです。

```js
assert.equal(fiveDice([4, 6, 2, 6, 5]), "pair");
```

`fiveDice([1, 4, 5, 6, 2])` は `"no pair"` を返すべきです。

```js
assert.equal(fiveDice([1, 4, 5, 6, 2]), "no pair");
```

`fiveDice([1, 3, 4, 6, 2])` は `"small straight"` を返すべきです。

```js
assert.equal(fiveDice([1, 3, 4, 6, 2]), "small straight");
```

`fiveDice([2, 2, 5, 2, 5])` は `"full house"` を返すべきです。

```js
assert.equal(fiveDice([2, 2, 5, 2, 5]), "full house");
```

`fiveDice([6, 4, 5, 6, 4])` は `"two pair"` を返すべきです。

```js
assert.equal(fiveDice([6, 4, 5, 6, 4]), "two pair");
```

# --seed--

## --seed-contents--

```js
function fiveDice(dice) {

  return dice;
}
```

# --solutions--

```js
function fiveDice(dice) {
  const counts = {};
  for (const d of dice) counts[d] = (counts[d] || 0) + 1;
  const values = Object.values(counts).sort((a, b) => b - a);
  const keys = Object.keys(counts).map(Number).sort((a, b) => a - b);

  const isLargeStraight = keys.length === 5 && keys[4] - keys[0] === 4;
  const isSmallStraight = !isLargeStraight && (() => {
    for (let i = 0; i <= 1; i++) {
      const slice = keys.slice(i, i + 4);
      if (slice.length === 4 && slice[3] - slice[0] === 3 &&
        new Set(slice).size === 4) return true;
    }
    return false;
  })();

  if (values[0] === 5) return "five of a kind";
  if (values[0] === 4) return "four of a kind";
  if (values[0] === 3 && values[1] === 2) return "full house";
  if (isLargeStraight) return "large straight";
  if (isSmallStraight) return "small straight";
  if (values[0] === 3) return "three of a kind";
  if (values[0] === 2 && values[1] === 2) return "two pair";
  if (values[0] === 2) return "pair";
  return "no pair";
}
```
