---
id: 6a1d9f98e819ed70a0e994df
title: "チャレンジ 333: イシュー トリアージ 2"
challengeType: 28
dashedName: challenge-333
---

# --description--

イシューのタイトルと現在のラベルの配列が与えられたら、以下のルールに基づいて更新されたラベルの配列を返してください。

イシューにラベルがない場合は、次を追加してください。

- タイトルに`"error"`または`"bug"`が含まれている場合は`"bug"`と`"needs triage"`を追加
- タイトルに`"feature"`または`"add"`が含まれている場合は`"enhancement"`と`"discussing"`を追加

それ以外の場合、与えられたラベルに次が含まれている場合は、

- `"needs triage"`がありタイトルに`"simple"`または`"easy"`が含まれている場合は`"needs triage"`を削除し`"good first issue"`を追加
- `"discussing"`がありタイトルに`"planned"`または`"next"`が含まれている場合は`"discussing"`を削除し`"on the roadmap"`を追加
- それ以外で`"needs triage"`または`"discussing"`があれば、それを削除して`"help wanted"`を追加

タイトルに次が含まれている場合は、

- `"security"`があれば`"critical"`ラベルを追加してください

# --hints--

`triageIssue("app crashes with error", [])`は`["bug", "needs triage"]`を返すべきです。

```js
assert.deepEqual(triageIssue("app crashes with error", []).sort(), ["bug", "needs triage"].sort());
```

`triageIssue("app crashes with error", ["bug", "needs triage"])`は`["bug", "help wanted"]`を返すべきです。

```js
assert.deepEqual(triageIssue("app crashes with error", ["bug", "needs triage"]).sort(), ["bug", "help wanted"].sort());
```

`triageIssue("add dark mode", [])`は`["enhancement", "discussing"]`を返すべきです。

```js
assert.deepEqual(triageIssue("add dark mode", []).sort(), ["enhancement", "discussing"].sort());
```

`triageIssue("add dark mode", ["enhancement", "discussing"])`は`["enhancement", "help wanted"]`を返すべきです。

```js
assert.deepEqual(triageIssue("add dark mode", ["enhancement", "discussing"]).sort(), ["enhancement", "help wanted"].sort());
```

`triageIssue("xss security bug", [])`は`["bug", "needs triage", "critical"]`を返すべきです。

```js
assert.deepEqual(triageIssue("xss security bug", []).sort(), ["bug", "needs triage", "critical"].sort());
```

`triageIssue("security vulnerability in auth", [])`は`["critical"]`を返すべきです。

```js
assert.deepEqual(triageIssue("security vulnerability in auth", []).sort(), ["critical"].sort());
```

`triageIssue("easy a11y fix", ["bug", "needs triage"])`は`["bug", "good first issue"]`を返すべきです。

```js
assert.deepEqual(triageIssue("easy a11y fix", ["bug", "needs triage"]).sort(), ["bug", "good first issue"].sort());
```

`triageIssue("planned api migration", ["enhancement", "discussing"])`は`["enhancement", "on the roadmap"]`を返すべきです。

```js
assert.deepEqual(triageIssue("planned api migration", ["enhancement", "discussing"]).sort(), ["enhancement", "on the roadmap"].sort());
```

`triageIssue("improve security", ["enhancement", "discussing"])`は`["enhancement", "help wanted", "critical"]`を返すべきです。

```js
assert.deepEqual(triageIssue("improve security", ["enhancement", "discussing"]).sort(), ["enhancement", "help wanted", "critical"].sort());
```

# --seed--

## --seed-contents--

```js
function triageIssue(title, labels) {

  return title;
}
```

# --solutions--

```js
function triageIssue(title, labels) {
  const result = [...labels];
  const t = title.toLowerCase();

  const has = label => result.includes(label);
  const add = label => { if (!has(label)) result.push(label); };
  const remove = label => { const i = result.indexOf(label); if (i > -1) result.splice(i, 1); };

  if (result.length === 0) {
    if (t.includes("error") || t.includes("bug")) {
      add("bug"); add("needs triage");
    } else if (t.includes("feature") || t.includes("add")) {
      add("enhancement"); add("discussing");
    }
  } else {
    if (has("needs triage") && (t.includes("simple") || t.includes("easy"))) {
      remove("needs triage"); add("good first issue");
    } else if (has("discussing") && (t.includes("planned") || t.includes("next"))) {
      remove("discussing"); add("on the roadmap");
    } else if (has("needs triage") || has("discussing")) {
      remove("needs triage"); remove("discussing"); add("help wanted");
    }
  }

  if (t.includes("security")) {
    add("critical");
  }

  return result;
}
```
