---
id: 6a1d9f98e819ed70a0e994d9
title: "チャレンジ 327: データベース マイグレーション"
challengeType: 28
dashedName: challenge-327
---

# --description--

2つのデータベースオブジェクトが与えられたとき、1つ目のオブジェクトにあって2つ目に欠けているプロパティを埋めた2つ目のオブジェクトを返してください。

- レコードにすでに存在するフィールドは上書きしないでください。

# --hints--

`migrateRecord({ username: "", posts: 0 }, { verified: true })` は `{ username: "", posts: 0, verified: true }` を返すべきです。

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { verified: true }), { username: "", posts: 0, verified: true });
```

`migrateRecord({ username: "", posts: 0 }, { username: "camper", posts: 5 })` は `{ username: "camper", posts: 5 }` を返すべきです。

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { username: "camper", posts: 5 }), { username: "camper", posts: 5 });
```

`migrateRecord({ username: "", posts: 0, verified: false }, { username: "camper" })` は `{ username: "camper", posts: 0, verified: false }` を返すべきです。

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0, verified: false }, { username: "camper" }), { username: "camper", posts: 0, verified: false });
```

`migrateRecord({ username: "", posts: 0 }, { username: "camper", role: "admin" })` は `{ username: "camper", role: "admin", posts: 0 }` を返すべきです。

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { username: "camper", role: "admin" }), { username: "camper", role: "admin", posts: 0 });
```

`migrateRecord({ username: "", email: "", posts: 0, verified: false, role: "user", banned: false }, { username: "camper", email: "camper@freecodecamp.org", role: "admin" })` は `{ username: "camper", email: "camper@freecodecamp.org", role: "admin", posts: 0, verified: false, banned: false }` を返すべきです。

```js
assert.deepEqual(migrateRecord({ username: "", email: "", posts: 0, verified: false, role: "user", banned: false }, { username: "camper", email: "camper@freecodecamp.org", role: "admin" }), { username: "camper", email: "camper@freecodecamp.org", role: "admin", posts: 0, verified: false, banned: false });
```

# --seed--

## --seed-contents--

```js
function migrateRecord(schema, record) {

  return schema;
}
```

# --solutions--

```js
function migrateRecord(schema, record) {
  return { ...schema, ...record };
}
```
