---
id: 6a1d9f98e819ed70a0e994dd
title: "チャレンジ 331: 最も近い倍数"
challengeType: 28
dashedName: challenge-331
---

# --description--

2つの整数が与えられたとき、最初の整数を2番目の整数の最も近い倍数に丸めてください。

# --hints--

`roundToNearestMultiple(5, 3)` は `6` を返すべきです。

```js
assert.equal(roundToNearestMultiple(5, 3), 6);
```

`roundToNearestMultiple(17, 4)` は `16` を返すべきです。

```js
assert.equal(roundToNearestMultiple(17, 4), 16);
```

`roundToNearestMultiple(43, 5)` は `45` を返すべきです。

```js
assert.equal(roundToNearestMultiple(43, 5), 45);
```

`roundToNearestMultiple(38, 11)` は `33` を返すべきです。

```js
assert.equal(roundToNearestMultiple(38, 11), 33);
```

`roundToNearestMultiple(93, 12)` は `96` を返すべきです。

```js
assert.equal(roundToNearestMultiple(93, 12), 96);
```

# --seed--

## --seed-contents--

```js
function roundToNearestMultiple(num, multiple) {

  return num;
}
```

# --solutions--

```js
function roundToNearestMultiple(num, multiple) {
  return Math.round(num / multiple) * multiple;
}
```
