---
id: 6a2037a68a0bc2aef0006002
title: "チャレンジ 340: ピッグラテン変換"
challengeType: 28
dashedName: challenge-340
---

# --description--

文字列が与えられたら、次のルールに従ってピッグラテンに変換してください。

- 単語が母音（`"a"`、`"e"`、`"i"`、`"o"`、または`"u"`）で始まる場合は、末尾に`"way"`を追加します。例えば、`"universe"`は`"universeway"`に変換されます。
- 単語が1つ以上の子音で始まる場合は、それらを末尾に移動し、`"ay"`を追加します。例えば、`"hello"`は`"ellohay"`に変換されます。
- 最初の文字の大文字・小文字は保持してください。例えば、`"Hello"`は`"Ellohay"`に変換されます。

# --hints--

`pigLatin("universe")`は`"universeway"`を返すべきです。

```js
assert.equal(pigLatin("universe"), "universeway");
```

`pigLatin("hello")`は`"ellohay"`を返すべきです。

```js
assert.equal(pigLatin("hello"), "ellohay");
```

`pigLatin("hello universe")`は`"ellohay universeway"`を返すべきです。

```js
assert.equal(pigLatin("hello universe"), "ellohay universeway");
```

`pigLatin("Hello universe")`は`"Ellohay universeway"`を返すべきです。

```js
assert.equal(pigLatin("Hello universe"), "Ellohay universeway");
```

`pigLatin("Pig Latin is fun")`は`"Igpay Atinlay isway unfay"`を返すべきです。

```js
assert.equal(pigLatin("Pig Latin is fun"), "Igpay Atinlay isway unfay");
```

`pigLatin("The quick brown fox jumped over the lazy dog")`は`"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"`を返すべきです。

```js
assert.equal(pigLatin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday");
```

# --seed--

## --seed-contents--

```js
function pigLatin(str) {

  return str;
}
```

# --solutions--

```js
function pigLatin(str) {
  const vowels = "aeiou";

  return str.split(" ").map(word => {
    const lowerWord = word.toLowerCase();

    if (vowels.includes(lowerWord[0])) {
      return word + "way";
    }

    let i = 0;
    while (i < word.length && !vowels.includes(lowerWord[i])) {
      i++;
    }

    const consonants = word.slice(0, i);
    const rest = word.slice(i);
    const result = rest + consonants + "ay";

    if (word[0] === word[0].toUpperCase()) {
      return result[0].toUpperCase() + result.slice(1).toLowerCase();
    }

    return result;
  }).join(" ");
}
```
