---
id: 6a2037a68a0bc2aef0006006
title: "チャレンジ 344: 黄金比"
challengeType: 28
dashedName: challenge-344
---

# --description--

2つの数字が与えられたとき、その比率が黄金比に近いかどうかを判定してください。

- 黄金比は `1.618` を使ってください
- 許容誤差は `0.01` としてください

# --hints--

`isGoldenRatio(21, 34)` は `true` を返すべきです。

```js
assert.isTrue(isGoldenRatio(21, 34));
```

`isGoldenRatio(15, 20)` は `false` を返すべきです。

```js
assert.isFalse(isGoldenRatio(15, 20));
```

`isGoldenRatio(8, 13)` は `true` を返すべきです。

```js
assert.isTrue(isGoldenRatio(8, 13));
```

`isGoldenRatio(10, 16)` は `false` を返すべきです。

```js
assert.isFalse(isGoldenRatio(10, 16));
```

`isGoldenRatio(1618, 1000)` は `true` を返すべきです。

```js
assert.isTrue(isGoldenRatio(1618, 1000));
```

`isGoldenRatio(88, 55)` は `false` を返すべきです。

```js
assert.isFalse(isGoldenRatio(88, 55));
```

# --seed--

## --seed-contents--

```js
function isGoldenRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function isGoldenRatio(a, b) {
  const goldenRatio = 1.618;
  const ratio = Math.max(a, b) / Math.min(a, b);
  return Math.abs(ratio - goldenRatio) <= 0.01;
}
```
