---
id: 6a2037a68a0bc2aef0006007
title: "チャレンジ 345: ワードブレンダー"
challengeType: 28
dashedName: challenge-345
---

# --description--

2つの単語が与えられたら、最初の単語の前半と2番目の単語の後半を組み合わせて新しい単語を返してください。

- 奇数文字の単語の場合、前半は短い方の半分になります。

# --hints--

`blendWords("turtle", "toucan")` は `"turcan"` を返すべきです。

```js
assert.equal(blendWords("turtle", "toucan"), "turcan");
```

`blendWords("chipmunk", "flamingo")` は `"chipingo"` を返すべきです。

```js
assert.equal(blendWords("chipmunk", "flamingo"), "chipingo");
```

`blendWords("falcon", "pelican")` は `"falican"` を返すべきです。

```js
assert.equal(blendWords("falcon", "pelican"), "falican");
```

`blendWords("hyena", "iguana")` は `"hyana"` を返すべきです。

```js
assert.equal(blendWords("hyena", "iguana"), "hyana");
```

`blendWords("scorpion", "gorilla")` は `"scorilla"` を返すべきです。

```js
assert.equal(blendWords("scorpion", "gorilla"), "scorilla");
```

`blendWords("platypus", "wolverine")` は `"platerine"` を返すべきです。

```js
assert.equal(blendWords("platypus", "wolverine"), "platerine");
```

# --seed--

## --seed-contents--

```js
function blendWords(word1, word2) {

  return word1;
}
```

# --solutions--

```js
function blendWords(word1, word2) {
  const firstHalf = word1.slice(0, Math.floor(word1.length / 2));
  const secondHalf = word2.slice(Math.floor(word2.length / 2));
  return firstHalf + secondHalf;
}
```
