---
id: 6a2037a68a0bc2aef0006003
title: "チャレンジ 341: 誕生日カウントダウン"
challengeType: 28
dashedName: challenge-341
---

# --description--

今日の日付と誕生日が与えられたとき、その人の次の誕生日までの日数を返してください。

- 今日の日付は先頭にゼロを付けた `"YYYY-MM-DD"` 形式の文字列で与えられます。例：`"2026-07-16"`。
- 誕生日は先頭にゼロを付けない `"M/D"` 形式の文字列で与えられます。例：`"9/7"`。
- 今日が誕生日の場合は、次の誕生日までの日数を返してください（`0` ではありません）。
- 閏年も考慮してください。

# --hints--

`daysUntilBirthday("2026-07-16", "9/7")` は `53` を返すべきです。

```js
assert.equal(daysUntilBirthday("2026-07-16", "9/7"), 53);
```

`daysUntilBirthday("2026-07-16", "3/22")` は `249` を返すべきです。

```js
assert.equal(daysUntilBirthday("2026-07-16", "3/22"), 249);
```

`daysUntilBirthday("2026-07-16", "7/16")` は `365` を返すべきです。

```js
assert.equal(daysUntilBirthday("2026-07-16", "7/16"), 365);
```

`daysUntilBirthday("2024-02-28", "3/1")` は `2` を返すべきです。

```js
assert.equal(daysUntilBirthday("2024-02-28", "3/1"), 2);
```

`daysUntilBirthday("2023-04-24", "12/30")` は `250` を返すべきです。

```js
assert.equal(daysUntilBirthday("2023-04-24", "12/30"), 250);
```

`daysUntilBirthday("2024-03-01", "2/29")` は `1460` を返すべきです。

```js
assert.equal(daysUntilBirthday("2024-03-01", "2/29"), 1460);
```

`daysUntilBirthday("2096-03-01", "2/29")` は `2920` を返すべきです。

```js
assert.equal(daysUntilBirthday("2096-03-01", "2/29"), 2920);
```

# --seed--

## --seed-contents--

```js
function daysUntilBirthday(today, birthday) {

  return today;
}
```

# --solutions--

```js
function daysUntilBirthday(today, birthday) {
  const [year, month, day] = today.split("-").map(Number);
  const [bMonth, bDay] = birthday.split("/").map(Number);

  const todayDate = new Date(year, month - 1, day);

  function isLeapYear(y) {
    return (y % 4 === 0 && y % 100 !== 0) || y % 400 === 0;
  }

  function nextValidBirthday(fromYear) {
    for (let y = fromYear; y <= fromYear + 8; y++) {
      if (bMonth === 2 && bDay === 29 && !isLeapYear(y)) continue;
      const date = new Date(y, bMonth - 1, bDay);
      if (date > todayDate) return date;
    }
  }

  const nextBirthday = nextValidBirthday(year);
  const diff = nextBirthday - todayDate;
  return Math.round(diff / (1000 * 60 * 60 * 24));
}
```
