---
id: 6a2037a68a0bc2aef0005fff
title: "チャレンジ 337: タリーカウンター"
challengeType: 28
dashedName: challenge-337
---

# --description--

タリーマークの文字列が与えられたら、表されている合計数を返してください。

- 各パイプ `"|"` は1つのカウントを表します。
- 5つ目のマークはスラッシュ `"/"` で表され、5つのグループを完成させます（`"||||/"`）。
- グループはスペースで区切られています。

# --hints--

`getTallyCount("||||")` は `4` を返すべきです。

```js
assert.equal(getTallyCount("||||"), 4);
```

`getTallyCount("||||/")` は `5` を返すべきです。

```js
assert.equal(getTallyCount("||||/"), 5);
```

`getTallyCount("||||/ |||")` は `8` を返すべきです。

```js
assert.equal(getTallyCount("||||/ |||"), 8);
```

`getTallyCount("||||/ ||||/ ||||/ ||")` は `17` を返すべきです。

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||"), 17);
```

`getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")` は `41` を返すべきです。

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41);
```

# --seed--

## --seed-contents--

```js
function getTallyCount(str) {

  return str;
}
```

# --solutions--

```js
function getTallyCount(str) {
  return str.replace(/ /g, "").length;
}
```
