---
id: 6a2037a68a0bc2aef0006000
title: "チャレンジ 338: ペット年齢計算機"
challengeType: 28
dashedName: challenge-338
---

# --description--

ペットの種類と人間の年齢が与えられたら、次の換算テーブルを使ってペットの年齢に換算した値を返してください。

| ペット | 乗数 |
|-----|-----------|
| `"dog"` | 7 |
| `"cat"` | 6 |
| `"rabbit"` | 8 |
| `"hamster"` | 30 |
| `"guinea pig"` | 12 |
| `"goldfish"` | 6 |
| `"bird"` | 5 |

# --hints--

`petYears("dog", 5)` は `35` を返すべきです。

```js
assert.equal(petYears("dog", 5), 35);
```

`petYears("cat", 9)` は `54` を返すべきです。

```js
assert.equal(petYears("cat", 9), 54);
```

`petYears("rabbit", 3)` は `24` を返すべきです。

```js
assert.equal(petYears("rabbit", 3), 24);
```

`petYears("hamster", 4)` は `120` を返すべきです。

```js
assert.equal(petYears("hamster", 4), 120);
```

`petYears("guinea pig", 5)` は `60` を返すべきです。

```js
assert.equal(petYears("guinea pig", 5), 60);
```

`petYears("goldfish", 2)` は `12` を返すべきです。

```js
assert.equal(petYears("goldfish", 2), 12);
```

`petYears("bird", 1)` は `5` を返すべきです。

```js
assert.equal(petYears("bird", 1), 5);
```

# --seed--

## --seed-contents--

```js
function petYears(pet, age) {

  return pet;
}
```

# --solutions--

```js
function petYears(pet, age) {
  const table = {
    dog: 7,
    cat: 6,
    rabbit: 8,
    hamster: 30,
    "guinea pig": 12,
    goldfish: 6,
    bird: 5,
  };

  return age * table[pet];
}
```
