---
id: 6a2037a68a0bc2aef0006001
title: "チャレンジ 339: 配列チャンク"
challengeType: 28
dashedName: challenge-339
---

# --description--

配列とチャンクサイズが与えられたら、そのサイズのサブ配列に分割した配列を返してください。

- 配列が均等に割り切れない場合、最後のチャンクは小さくなることがあります。

# --hints--

`chunkArray([1, 2, 3, 4, 5, 6], 3)` は `[[1, 2, 3], [4, 5, 6]]` を返すべきです。

```js
assert.deepEqual(chunkArray([1, 2, 3, 4, 5, 6], 3), [[1, 2, 3], [4, 5, 6]]);
```

`chunkArray([1, "two", 3, "four", 5, "six", 7, "eight"], 2)` は `[[1, "two"], [3, "four"], [5, "six"], [7, "eight"]]` を返すべきです。

```js
assert.deepEqual(chunkArray([1, "two", 3, "four", 5, "six", 7, "eight"], 2), [[1, "two"], [3, "four"], [5, "six"], [7, "eight"]]);
```

`chunkArray([1, 2, 3, 4, 5], 3)` は `[[1, 2, 3], [4, 5]]` を返すべきです。

```js
assert.deepEqual(chunkArray([1, 2, 3, 4, 5], 3), [[1, 2, 3], [4, 5]]);
```

`chunkArray(["a", "b", "c", "d", "e"], 1)` は `[["a"], ["b"], ["c"], ["d"], ["e"]]` を返すべきです。

```js
assert.deepEqual(chunkArray(["a", "b", "c", "d", "e"], 1), [["a"], ["b"], ["c"], ["d"], ["e"]]);
```

`chunkArray([1, 2, 3], 5)` は `[[1, 2, 3]]` を返すべきです。

```js
assert.deepEqual(chunkArray([1, 2, 3], 5), [[1, 2, 3]]);
```

# --seed--

## --seed-contents--

```js
function chunkArray(arr, size) {

  return arr;
}
```

# --solutions--

```js
function chunkArray(arr, size) {
  const result = [];
  for (let i = 0; i < arr.length; i += size) {
    result.push(arr.slice(i, i + size));
  }
  return result;
}
```
