---
id: 6a22d77ddf034bc4e35b1d5b
title: "チャレンジ 354: コントラスト評価 3"
challengeType: 28
dashedName: challenge-354
---

# --description--

RGB値を表す2つの配列とテキストが大きいかどうかを示すブール値が与えられたら、次の方法でWCAGのコントラスト評価を返してください。

まず、各RGB値を相対輝度に変換します。

- 各チャンネル`[R, G, B]`を255で割り、`0`から`1`の範囲の値にします
- 各チャンネルにガンマ補正の式を適用します:
  - チャンネル値が`0.04045`以下の場合: `channel / 12.92`
  - それ以外の場合: `((channel + 0.055) / 1.055) ^ 2.4`
- 輝度を計算します: `0.2126 * R + 0.7152 * G + 0.0722 * B`

次に、各輝度値に0.05を加え、明るい方の値を暗い方の値で割ってコントラスト比を計算します。明るい方は常に最初の引数になります。

次のテーブルを使ってコントラスト比に基づく評価を返してください。

| 評価 | 通常テキスト | 大きなテキスト |
|--------|-------------|------------|
| `"AAA"` | 7.0以上 | 4.5以上 |
| `"AA"` | 4.5以上 | 3.0以上 |
| `"Fail"` | 4.5未満 | 3.0未満 |

# --hints--

`getContrastRating([255, 255, 255], [0, 0, 0], false)` は `"AAA"` を返すべきです。

```js
assert.equal(getContrastRating([255, 255, 255], [0, 0, 0], false), "AAA");
```

`getContrastRating([215, 188, 188], [55, 55, 55], false)` は `"AA"` を返すべきです。

```js
assert.equal(getContrastRating([215, 188, 188], [55, 55, 55], false), "AA");
```

`getContrastRating([143, 144, 210], [46, 47, 61], false)` は `"Fail"` を返すべきです。

```js
assert.equal(getContrastRating([143, 144, 210], [46, 47, 61], false), "Fail");
```

`getContrastRating([167, 167, 210], [53, 10, 53], true)` は `"AAA"` を返すべきです。

```js
assert.equal(getContrastRating([167, 167, 210], [53, 10, 53], true), "AAA");
```

`getContrastRating([135, 147, 155], [60, 70, 90], true)` は `"AA"` を返すべきです。

```js
assert.equal(getContrastRating([135, 147, 155], [60, 70, 90], true), "AA");
```

`getContrastRating([125, 210, 195], [105, 130, 90], true)` は `"Fail"` を返すべきです。

```js
assert.equal(getContrastRating([125, 210, 195], [105, 130, 90], true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {

  return rgb1;
}
```

# --solutions--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {
  function toLuminance([r, g, b]) {
    return [r, g, b].map(channel => {
      channel = channel / 255;
      return channel <= 0.04045
        ? channel / 12.92
        : ((channel + 0.055) / 1.055) ** 2.4;
    }).reduce((sum, c, i) => sum + c * [0.2126, 0.7152, 0.0722][i], 0);
  }

  const l1 = toLuminance(rgb1);
  const l2 = toLuminance(rgb2);
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
