---
id: 6a22d77ddf034bc4e35b1d58
title: "チャレンジ 351: プロニック数"
challengeType: 28
dashedName: challenge-351
---

# --description--

数値がプロニック数かどうかを判定してください。

プロニック数とは、連続する2つの整数の積です。例えば、6は2 * 3 = 6なのでプロニック数です。

# --hints--

`isPronic(6)` は `true` を返すべきです。

```js
assert.isTrue(isPronic(6));
```

`isPronic(15)` は `false` を返すべきです。

```js
assert.isFalse(isPronic(15));
```

`isPronic(12)` は `true` を返すべきです。

```js
assert.isTrue(isPronic(12));
```

`isPronic(132)` は `true` を返すべきです。

```js
assert.isTrue(isPronic(132));
```

`isPronic(80)` は `false` を返すべきです。

```js
assert.isFalse(isPronic(80));
```

`isPronic(0)` は `true` を返すべきです。

```js
assert.isTrue(isPronic(0));
```

# --seed--

## --seed-contents--

```js
function isPronic(n) {

  return n;
}
```

# --solutions--

```js
function isPronic(n) {
  const k = Math.floor(Math.sqrt(n));
  return k * (k + 1) === n;
}
```
