---
id: 6a22d77ddf034bc4e35b1d55
title: "チャレンジ 348: ローン計算機"
challengeType: 28
dashedName: challenge-348
---

# --description--

ローンの元本、年利率のパーセンテージ、固定の月々の支払い額が与えられたとき、ローンが完済されるまでの各月の支払い後の残高を配列で返してください。

- 毎月、残高に対して月利率 `(annual rate / 100) / 12` を使って利息を計算し、その後月々の支払い額を差し引きます。
- 各残高は最も近いドルに丸めて返してください。
- 返す配列にはローンの元本も含めてください。配列の最初の要素は常にローンの元本で、最後の要素は常に0になります。

# --hints--

`getLoanSchedule(1000, 0, 200)` は `[1000, 800, 600, 400, 200, 0]` を返すべきです。

```js
assert.deepEqual(getLoanSchedule(1000, 0, 200), [1000, 800, 600, 400, 200, 0]);
```

`getLoanSchedule(1000, 5, 200)` は `[1000, 804, 608, 410, 212, 13, 0]` を返すべきです。

```js
assert.deepEqual(getLoanSchedule(1000, 5, 200), [1000, 804, 608, 410, 212, 13, 0]);
```

`getLoanSchedule(10, 50, 1)` は `[10, 9, 9, 8, 8, 7, 6, 5, 5, 4, 3, 2, 1, 0, 0]` を返すべきです。

```js
assert.deepEqual(getLoanSchedule(10, 50, 1), [10, 9, 9, 8, 8, 7, 6, 5, 5, 4, 3, 2, 1, 0, 0]);
```

`getLoanSchedule(5500, 8, 400)` は `[5500, 5137, 4771, 4403, 4032, 3659, 3283, 2905, 2525, 2141, 1756, 1367, 977, 583, 187, 0]` を返すべきです。

```js
assert.deepEqual(getLoanSchedule(5500, 8, 400), [5500, 5137, 4771, 4403, 4032, 3659, 3283, 2905, 2525, 2141, 1756, 1367, 977, 583, 187, 0]);
```

`getLoanSchedule(50000, 5.2, 1650)` は `[50000, 48567, 47127, 45681, 44229, 42771, 41306, 39835, 38358, 36874, 35384, 33887, 32384, 30874, 29358, 27835, 26306, 24770, 23227, 21678, 20122, 18559, 16990, 15413, 13830, 12240, 10643, 9039, 7428, 5810, 4186, 2554, 915, 0]` を返すべきです。

```js
assert.deepEqual(getLoanSchedule(50000, 5.2, 1650), [50000, 48567, 47127, 45681, 44229, 42771, 41306, 39835, 38358, 36874, 35384, 33887, 32384, 30874, 29358, 27835, 26306, 24770, 23227, 21678, 20122, 18559, 16990, 15413, 13830, 12240, 10643, 9039, 7428, 5810, 4186, 2554, 915, 0]);
```

# --seed--

## --seed-contents--

```js
function getLoanSchedule(loanAmount, annualRate, monthlyPayment) {

  return loanAmount;
}
```

# --solutions--

```js
function getLoanSchedule(loanAmount, annualRate, monthlyPayment) {
  const monthlyRate = (annualRate / 100) / 12;
  const balances = [loanAmount];
  let balance = loanAmount;

  while (balance > 0) {
    balance += balance * monthlyRate;
    balance -= monthlyPayment;
    balances.push(Math.round(Math.max(balance, 0)));
  }

  return balances;
}
```
