---
id: 6a22d77ddf034bc4e35b1d54
title: "チャレンジ 347: ゲーム理論"
challengeType: 28
dashedName: challenge-347
---

# --description--

2人のプレイヤーの戦略を表す同じ長さの2つの文字列が与えられたとき、スコアを配列 `[player1, player2]` として返してください。

- 与えられた文字列は、`"C"`（協力）か`"D"`（裏切り）のいずれかの文字のみを含みます。
- 各文字は1ラウンドを表し、スコアは次のように計算されます：
  - 両方のプレイヤーが協力した場合、各自3点。
  - 両方のプレイヤーが裏切った場合、各自1点。
  - 片方が裏切り、もう片方が協力した場合、裏切ったプレイヤーは5点、協力したプレイヤーは0点。

# --hints--

`playGame("CCCC", "CCCC")` は `[12, 12]` を返すべきです。

```js
assert.deepEqual(playGame("CCCC", "CCCC"), [12, 12]);
```

`playGame("DDDD", "DDDD")` は `[4, 4]` を返すべきです。

```js
assert.deepEqual(playGame("DDDD", "DDDD"), [4, 4]);
```

`playGame("CCDD", "CDDD")` は `[5, 10]` を返すべきです。

```js
assert.deepEqual(playGame("CCDD", "CDDD"), [5, 10]);
```

`playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD")` は `[24, 34]` を返すべきです。

```js
assert.deepEqual(playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34]);
```

`playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` は `[66, 21]` を返すべきです。

```js
assert.deepEqual(playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21]);
```

# --seed--

## --seed-contents--

```js
function playGame(p1, p2) {

  return p1;
}
```

# --solutions--

```js
function playGame(p1, p2) {
  let scores = [0, 0];

  for (let i = 0; i < p1.length; i++) {
    if (p1[i] === 'C' && p2[i] === 'C') {
      scores[0] += 3; scores[1] += 3;
    } else if (p1[i] === 'D' && p2[i] === 'D') {
      scores[0] += 1; scores[1] += 1;
    } else if (p1[i] === 'D' && p2[i] === 'C') {
      scores[0] += 5; scores[1] += 0;
    } else {
      scores[0] += 0; scores[1] += 5;
    }
  }

  return scores;
}
```
