---
id: 6a22d77ddf034bc4e35b1d5a
title: "チャレンジ 353: コントラスト評価 2"
challengeType: 28
dashedName: challenge-353
---

# --description--

2つの相対輝度値とテキストが大きいかどうかを示すブール値が与えられたら、次の方法でWCAGのコントラスト評価を返してください。

各輝度値に0.05を加え、明るい方の値を暗い方の値で割ってコントラスト比を計算します。明るい方の値は常に最初の引数になります。

次のテーブルを使って、コントラスト比に基づく評価を返してください。

| 評価 | 通常のテキスト | 大きなテキスト |
|--------|-------------|------------|
| `"AAA"` | 7.0以上 | 4.5以上 |
| `"AA"` | 4.5以上 | 3.0以上 |
| `"Fail"` | 4.5未満 | 3.0未満 |

# --hints--

`getContrastRating(1.0, 0.0, false)` は `"AAA"` を返すべきです。

```js
assert.equal(getContrastRating(1.0, 0.0, false), "AAA");
```

`getContrastRating(0.9015, 0.1364, false)` は `"AA"` を返すべきです。

```js
assert.equal(getContrastRating(0.9015, 0.1364, false), "AA");
```

`getContrastRating(0.8965, 0.1628, false)` は `"Fail"` を返すべきです。

```js
assert.equal(getContrastRating(0.8965, 0.1628, false), "Fail");
```

`getContrastRating(0.7469, 0.0957, true)` は `"AAA"` を返すべきです。

```js
assert.equal(getContrastRating(0.7469, 0.0957, true), "AAA");
```

`getContrastRating(0.7489, 0.2018, true)` は `"AA"` を返すべきです。

```js
assert.equal(getContrastRating(0.7489, 0.2018, true), "AA");
```

`getContrastRating(0.6571, 0.1974, true)` は `"Fail"` を返すべきです。

```js
assert.equal(getContrastRating(0.6571, 0.1974, true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(l1, l2, isLargeText) {

  return l1;
}
```

# --solutions--

```js
function getContrastRating(l1, l2, isLargeText) {
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
