---
id: 6a22d77ddf034bc4e35b1d59
title: "チャレンジ 352: コントラスト評価 1"
challengeType: 28
dashedName: challenge-352
---

# --description--

コントラスト比とテキストが大きいかどうかを示すブール値が与えられたら、次のテーブルを使ってWCAG評価を返してください。

| 評価 | 通常のテキスト | 大きなテキスト |
|--------|-------------|------------|
| `"AAA"` | 7.0以上 | 4.5以上 |
| `"AA"` | 4.5以上 | 3.0以上 |
| `"Fail"` | 4.5未満 | 3.0未満 |

# --hints--

`getContrastRating("7.5", false)` は `"AAA"` を返すべきです。

```js
assert.equal(getContrastRating("7.5", false), "AAA");
```

`getContrastRating("4.8", false)` は `"AA"` を返すべきです。

```js
assert.equal(getContrastRating("4.8", false), "AA");
```

`getContrastRating("4.2", false)` は `"Fail"` を返すべきです。

```js
assert.equal(getContrastRating("4.2", false), "Fail");
```

`getContrastRating("4.5", true)` は `"AAA"` を返すべきです。

```js
assert.equal(getContrastRating("4.5", true), "AAA");
```

`getContrastRating("3.0", true)` は `"AA"` を返すべきです。

```js
assert.equal(getContrastRating("3.0", true), "AA");
```

`getContrastRating("2.7", false)` は `"Fail"` を返すべきです。

```js
assert.equal(getContrastRating("2.7", false), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(ratio, isLargeText) {

  return ratio;
}
```

# --solutions--

```js
function getContrastRating(ratio, isLargeText) {
  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
