---
id: 6a22d77ddf034bc4e35b1d57
title: "チャレンジ 350: 文字の距離"
challengeType: 28
dashedName: challenge-350
---

# --description--

長さが同じ2つの文字列が与えられたとき、それぞれの文字のペア間の最短距離の合計を返してください。

- 入力は小文字のみを含みます
- アルファベットは円環として扱われるため、`a` と `z` の距離は1です。

# --hints--

`letterDistance("abc", "bcd")` は `3` を返すべきです。

```js
assert.equal(letterDistance("abc", "bcd"), 3);
```

`letterDistance("abc", "xyz")` は `9` を返すべきです。

```js
assert.equal(letterDistance("abc", "xyz"), 9);
```

`letterDistance("encrypt", "decrypt")` は `10` を返すべきです。

```js
assert.equal(letterDistance("encrypt", "decrypt"), 10);
```

`letterDistance("algorithm", "codeblock")` は `43` を返すべきです。

```js
assert.equal(letterDistance("algorithm", "codeblock"), 43);
```

`letterDistance("lobster", "penguin")` は `47` を返すべきです。

```js
assert.equal(letterDistance("lobster", "penguin"), 47);
```

`letterDistance("alligator", "crocodile")` は `55` を返すべきです。

```js
assert.equal(letterDistance("alligator", "crocodile"), 55);
```

# --seed--

## --seed-contents--

```js
function letterDistance(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function letterDistance(str1, str2) {
  return str1.split("").reduce((sum, char, i) => {
    const diff = Math.abs(char.charCodeAt(0) - str2[i].charCodeAt(0));
    return sum + Math.min(diff, 26 - diff);
  }, 0);
}
```
