---
id: 6a22d77ddf034bc4e35b1d56
title: "チャレンジ 349: セル信号"
challengeType: 28
dashedName: challenge-349
---

# --description--

3つのセルタワーの読み取り値が含まれるグリッドが与えられたとき、電話の位置を特定してください。

- グリッドの各セルは `0`（タワーなし）か、電話までのセル数を表す正の整数で、直線（水平、垂直、または斜め）で測定されます。
- 3つのタワーすべてから正しいセル数だけ離れているセルの `[row, col]` を返してください。
- 解は必ず1つだけ存在します。

# --hints--

`findSignal([[0, 0, 1], [0, 1, 0], [0, 0, 1]])` は `[1, 2]` を返すべきです。

```js
assert.deepEqual(findSignal([[0, 0, 1], [0, 1, 0], [0, 0, 1]]), [1, 2]);
```

`findSignal([[0, 2, 0], [1, 0, 0], [0, 0, 1]])` は `[2, 1]` を返すべきです。

```js
assert.deepEqual(findSignal([[0, 2, 0], [1, 0, 0], [0, 0, 1]]), [2, 1]);
```

`findSignal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]])` は `[2, 2]` を返すべきです。

```js
assert.deepEqual(findSignal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]]), [2, 2]);
```

`findSignal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]])` は `[3, 4]` を返すべきです。

```js
assert.deepEqual(findSignal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]]), [3, 4]);
```

`findSignal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]])` は `[3, 3]` を返すべきです。

```js
assert.deepEqual(findSignal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]]), [3, 3]);
```

# --seed--

## --seed-contents--

```js
function findSignal(grid) {

  return grid;
}
```

# --solutions--

```js
function findSignal(grid) {
  const rows = grid.length;
  const cols = grid[0].length;

  const towers = [];
  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (grid[r][c] !== 0) towers.push([r, c, grid[r][c]]);
    }
  }

  function isValid(r, c, tr, tc, dist) {
    const dr = Math.abs(r - tr);
    const dc = Math.abs(c - tc);
    return (dr === dist || dc === dist) && (dr === 0 || dc === 0 || dr === dc);
  }

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (grid[r][c] !== 0) continue;
      if (towers.every(([tr, tc, dist]) => isValid(r, c, tr, tc, dist))) {
        return [r, c];
      }
    }
  }
}
```
