---
id: 6a26df95efa55a2524399743
title: "チャレンジ 362: ノノグラムバリデーター"
challengeType: 28
dashedName: challenge-362
---

# --description--

手がかりの数字の配列とセルの配列が与えられたとき、そのセルがノノグラムの手がかりを満たしているかどうかを判定してください。

- 手がかりは、連続した塗りつぶされたセルの長さを順番に表す数字の配列です。例えば、手がかりが`[3, 2]`の場合、3つの連続した塗りつぶしセルの後に少なくとも1つの空セルを挟んで2つの連続した塗りつぶしセルがあることを意味します。
- 行は1（塗りつぶし）と0（空）の配列です。

# --hints--

`isValidNonogram([3, 2], [1, 1, 1, 0, 1, 1])`は`true`を返すべきです。

```js
assert.isTrue(isValidNonogram([3, 2], [1, 1, 1, 0, 1, 1]));
```

`isValidNonogram([3, 2], [0, 1, 1, 1, 1, 1])`は`false`を返すべきです。

```js
assert.isFalse(isValidNonogram([3, 2], [0, 1, 1, 1, 1, 1]));
```

`isValidNonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1])`は`false`を返すべきです。

```js
assert.isFalse(isValidNonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1]));
```

`isValidNonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0])`は`true`を返すべきです。

```js
assert.isTrue(isValidNonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0]));
```

`isValidNonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0])`は`true`を返すべきです。

```js
assert.isTrue(isValidNonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0]));
```

`isValidNonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0])`は`false`を返すべきです。

```js
assert.isFalse(isValidNonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0]));
```

# --seed--

## --seed-contents--

```js
function isValidNonogram(clue, cells) {

  return clue;
}
```

# --solutions--

```js
function isValidNonogram(clue, cells) {
  const runs = [];
  let count = 0;

  for (let i = 0; i <= cells.length; i++) {
    if (cells[i] === 1) {
      count++;
    } else if (count > 0) {
      runs.push(count);
      count = 0;
    }
  }

  if (runs.length !== clue.length) return false;
  return runs.every((run, i) => run === clue[i]);
}
```
