---
id: 6a26df3e2988bcdded204897
title: "チャレンジ 361: 話し言葉の時間"
challengeType: 28
dashedName: challenge-361
---

# --description--

アナログ時計の時針と分針の角度（12時から時計回りの度数）が与えられたら、話し言葉の英語で時間を返してください。

分針の角度を分に変換します（360°＝60分）。その後、以下のルールを使ってください。

| 分 | 話し言葉 |
|---------|--------|
| 0 | `"Y o'clock"` |
| 15 | `"quarter past Y"` |
| 1–29（15を除く） | `"X minutes past Y"` |
| 30 | `"half past Y"` |
| 45 | `"quarter to Z"` |
| 31–59（45を除く） | `"X minutes to Z"`（ここで`X`は60 - 分）|

`Y`は現在の時、`Z`は次の時で、どちらも時針の角度（360°＝12時間）から導き出されます。

注意：針の角度は正確に数字に合わない場合があるので、何らかの方法で丸めることを検討してください。

# --hints--

`getSpokenTime(90, 0)`は`"3 o'clock"`を返すべきです。

```js
assert.equal(getSpokenTime(90, 0), "3 o'clock");
```

`getSpokenTime(160, 120)`は`"20 minutes past 5"`を返すべきです。

```js
assert.equal(getSpokenTime(160, 120), "20 minutes past 5");
```

`getSpokenTime(255, 180)`は`"half past 8"`を返すべきです。

```js
assert.equal(getSpokenTime(255, 180), "half past 8");
```

`getSpokenTime(67.5, 92)`は`"quarter past 2"`を返すべきです。

```js
assert.equal(getSpokenTime(67.5, 92), "quarter past 2");
```

`getSpokenTime(200, 240)`は`"20 minutes to 7"`を返すべきです。

```js
assert.equal(getSpokenTime(200, 240), "20 minutes to 7");
```

`getSpokenTime(322.5, 273)`は`"quarter to 11"`を返すべきです。

```js
assert.equal(getSpokenTime(322.5, 273), "quarter to 11");
```

`getSpokenTime(117.5, 335)`は`"5 minutes to 4"`を返すべきです。

```js
assert.equal(getSpokenTime(117.5, 335), "5 minutes to 4");
```

# --seed--

## --seed-contents--

```js
function getSpokenTime(hourAngle, minuteAngle) {

  return hourAngle;
}
```

# --solutions--

```js
function getSpokenTime(hourAngle, minuteAngle) {
  const minutes = Math.floor(minuteAngle / 6);
  const hour = Math.floor(hourAngle / 30) || 12;
  const nextHour = (hour % 12) + 1;

  if (minutes === 0) return `${hour} o'clock`;
  if (minutes === 15) return `quarter past ${hour}`;
  if (minutes === 30) return `half past ${hour}`;
  if (minutes === 45) return `quarter to ${nextHour}`;
  if (minutes < 30) return `${minutes} minutes past ${hour}`;
  return `${60 - minutes} minutes to ${nextHour}`;
}
```
