---
id: 6a26df3e2988bcdded204893
title: "チャレンジ 357: 食物連鎖"
challengeType: 28
dashedName: challenge-357
---

# --description--

`[predator, prey]` のペアの配列が与えられたら、頂点捕食者から最下層までの食物連鎖を返してください。

- 頂点捕食者とは、他の動物に捕食されることがない動物です。
- 連鎖は文字列の配列として返してください。

# --hints--

`getFoodChain([["cat", "mouse"]])` は `["cat", "mouse"]` を返すべきです。

```js
assert.deepEqual(getFoodChain([["cat", "mouse"]]), ["cat", "mouse"]);
```

`getFoodChain([["wolf", "deer"], ["deer", "grass"]])` は `["wolf", "deer", "grass"]` を返すべきです。

```js
assert.deepEqual(getFoodChain([["wolf", "deer"], ["deer", "grass"]]), ["wolf", "deer", "grass"]);
```

`getFoodChain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]])` は `["hawk", "snake", "frog", "fly"]` を返すべきです。

```js
assert.deepEqual(getFoodChain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]]), ["hawk", "snake", "frog", "fly"]);
```

`getFoodChain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]])` は `["eagle", "fox", "rabbit", "grass"]` を返すべきです。

```js
assert.deepEqual(getFoodChain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]]), ["eagle", "fox", "rabbit", "grass"]);
```

`getFoodChain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]])` は `["orca", "seal", "salmon", "herring", "shrimp", "plankton"]` を返すべきです。

```js
assert.deepEqual(getFoodChain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]]), ["orca", "seal", "salmon", "herring", "shrimp", "plankton"]);
```

# --seed--

## --seed-contents--

```js
function getFoodChain(pairs) {

  return pairs;
}
```

# --solutions--

```js
function getFoodChain(pairs) {
  const prey = new Set(pairs.map(([, p]) => p));
  const map = Object.fromEntries(pairs);
  const apex = pairs.map(([p]) => p).find(p => !prey.has(p));
  const chain = [apex];
  while (map[chain.at(-1)]) chain.push(map[chain.at(-1)]);
  return chain;
}
```
