---
id: 6a26df3e2988bcdded204896
title: "チャレンジ 360: 話し言葉の時間表現"
challengeType: 28
dashedName: challenge-360
---

# --description--

秒数が与えられたら、話し言葉の英語で時間の長さを返してください。

- 時間の長さを時間、分、秒に分けてください。
- 0の値は省略してください。
- 単数形か複数形かを適切に使い分けてください（`"1 hour"`、`"2 hours"`）。
- もしあれば、最後の2つの単位は`"and"`でつなぎ、最後から2番目と3番目の単位はカンマで区切ってください（`"1 hour, 2 minutes and 3 seconds"`）。

# --hints--

`getSpokenDuration(3723)` は `"1 hour, 2 minutes and 3 seconds"` を返すべきです。

```js
assert.equal(getSpokenDuration(3723), "1 hour, 2 minutes and 3 seconds");
```

`getSpokenDuration(7295)` は `"2 hours, 1 minute and 35 seconds"` を返すべきです。

```js
assert.equal(getSpokenDuration(7295), "2 hours, 1 minute and 35 seconds");
```

`getSpokenDuration(8521)` は `"2 hours, 22 minutes and 1 second"` を返すべきです。

```js
assert.equal(getSpokenDuration(8521), "2 hours, 22 minutes and 1 second");
```

`getSpokenDuration(435)` は `"7 minutes and 15 seconds"` を返すべきです。

```js
assert.equal(getSpokenDuration(435), "7 minutes and 15 seconds");
```

`getSpokenDuration(14455)` は `"4 hours and 55 seconds"` を返すべきです。

```js
assert.equal(getSpokenDuration(14455), "4 hours and 55 seconds");
```

`getSpokenDuration(72000)` は `"20 hours"` を返すべきです。

```js
assert.equal(getSpokenDuration(72000), "20 hours");
```

`getSpokenDuration(1)` は `"1 second"` を返すべきです。

```js
assert.equal(getSpokenDuration(1), "1 second");
```

# --seed--

## --seed-contents--

```js
function getSpokenDuration(seconds) {

  return seconds;
}
```

# --solutions--

```js
function getSpokenDuration(seconds) {
  const hours = Math.floor(seconds / 3600);
  const minutes = Math.floor((seconds % 3600) / 60);
  const secs = seconds % 60;

  const parts = [];
  if (hours) parts.push(`${hours} ${hours === 1 ? "hour" : "hours"}`);
  if (minutes) parts.push(`${minutes} ${minutes === 1 ? "minute" : "minutes"}`);
  if (secs) parts.push(`${secs} ${secs === 1 ? "second" : "seconds"}`);

  if (parts.length === 1) return parts[0];
  if (parts.length === 2) return `${parts[0]} and ${parts[1]}`;
  return `${parts[0]}, ${parts[1]} and ${parts[2]}`;
}
```
