---
id: 6a26df3e2988bcdded204894
title: "チャレンジ 358: 絵文字翻訳"
challengeType: 28
dashedName: challenge-358
---

# --description--

絵文字の文字列が与えられたら、次のテーブルを使ってフレーズを返してください。

| Emoji | Word |
|-------|------|
| 👶 | `"baby"` |
| 🐱 | `"cat"` |
| 🐕 | `"dog"` |
| 🐟 | `"fish"` |
| 🥵 | `"hot"` |
| 🧊 | `"ice"` |
| 🪨 | `"rock"` |
| 🦈 | `"shark"` |
| 🍲 | `"soup"` |
| ⭐ | `"star"` |

単語はスペースで区切って返してください。

# --hints--

`getEmojiPhrase("🪨⭐")` は `"rock star"` を返すべきです。

```js
assert.equal(getEmojiPhrase("🪨⭐"), "rock star");
```

`getEmojiPhrase("🥵🐕")` は `"hot dog"` を返すべきです。

```js
assert.equal(getEmojiPhrase("🥵🐕"), "hot dog");
```

`getEmojiPhrase("👶🦈")` は `"baby shark"` を返すべきです。

```js
assert.equal(getEmojiPhrase("👶🦈"), "baby shark");
```

`getEmojiPhrase("⭐🐟")` は `"star fish"` を返すべきです。

```js
assert.equal(getEmojiPhrase("⭐🐟"), "star fish");
```

`getEmojiPhrase("🧊🧊👶")` は `"ice ice baby"` を返すべきです。

```js
assert.equal(getEmojiPhrase("🧊🧊👶"), "ice ice baby");
```

`getEmojiPhrase("🐱🐟🍲")` は `"cat fish soup"` を返すべきです。

```js
assert.equal(getEmojiPhrase("🐱🐟🍲"), "cat fish soup");
```

# --seed--

## --seed-contents--

```js
function getEmojiPhrase(str) {

  return str;
}
```

# --solutions--

```js
function getEmojiPhrase(str) {
  const table = {
    '👶': 'baby',
    '🐱': 'cat',
    '🐕': 'dog',
    '🐟': 'fish',
    '🥵': 'hot',
    '🧊': 'ice',
    '🪨': 'rock',
    '🦈': 'shark',
    '🍲': 'soup',
    '⭐': 'star',
  };

  return [...str].map(emoji => table[emoji]).join(' ');
}
```
