---
id: 6a26df95efa55a2524399745
title: "チャレンジ 364: 2つのバケツの間で"
challengeType: 28
dashedName: challenge-364
---

# --description--

RGBカラーと満杯度を持つ2つのペンキのバケツが与えられたとき、混ぜたRGBカラーを3つの整数の配列として返してください。

- 各バケツは、`color`プロパティ（3つの整数の配列`[r, g, b]`）と`fullness`プロパティ（0～100）を持つオブジェクト（JavaScript）または辞書（Python）です。
- 混ぜた色は、満杯度に基づく2つの色の各チャンネルの加重平均で、各チャンネルは最も近い整数に丸められます。

# --hints--

`mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 })` は `[125, 125, 125]` を返すべきです。

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 }), [125, 125, 125]);
```

`mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 })` は `[200, 200, 200]` を返すべきです。

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 }), [200, 200, 200]);
```

`mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 })` は `[100, 150, 200]` を返すべきです。

```js
assert.deepEqual(mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 }), [100, 150, 200]);
```

`mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 })` は `[114, 184, 170]` を返すべきです。

```js
assert.deepEqual(mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 }), [114, 184, 170]);
```

`mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 })` は `[68, 163, 123]` を返すべきです。

```js
assert.deepEqual(mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 }), [68, 163, 123]);
```

# --seed--

## --seed-contents--

```js
function mixPaint(bucket1, bucket2) {

  return bucket1;
}
```

# --solutions--

```js
function mixPaint(bucket1, bucket2) {
  const total = bucket1.fullness + bucket2.fullness;
  return bucket1.color.map((channel, i) =>
    Math.round((channel * bucket1.fullness + bucket2.color[i] * bucket2.fullness) / total)
  );
}
```
