---
id: 616965351e74d4689eb6de30
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Sie können mehrere `meta`-Elemente auf einer Webseite haben. Jedes `meta`-Element fügt Informationen über die Seite hinzu, die von anderen HTML-Elementen nicht ausgedrückt werden können.

Fügen Sie ein weiteres `meta`-Element innerhalb des `head` ein. Geben Sie ihm ein `name`-Attribut mit dem Wert `"viewport"` und ein `content`-Attribut mit dem Wert `"width=device-width, initial-scale=1.0"`, damit Ihre Seite auf allen Geräten gleich aussieht.

# --hints--

Sie sollten zwei `meta`-Elemente haben.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 2);
```

Ihr neues `meta`-Element sollte ein Void-Element sein, es darf kein schließendes Tag `</meta>` haben.

```js
assert.notMatch(code, /<\/meta>/i);
```


Ihr neues `meta`-Element sollte ein `name`-Attribut mit dem Wert `"viewport"` und ein `content`-Attribut mit dem Wert `"width=device-width, initial-scale=1.0"` haben.

```js
const meta = [...document.querySelectorAll('meta')];
const contentAttrRegex = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/i;
const target = meta?.find(m => m?.getAttribute('name')?.toLowerCase()?.replace(/\s*/g, '') === 'viewport' && contentAttrRegex.test(m?.getAttribute('content')));
assert.exists(target);
```

Ihre zwei `meta`-Elemente sollten sich innerhalb des `head`-Elements befinden.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('meta').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
