---
id: 61695c4aad56f95497c19583
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Denken Sie daran, dass das `title`-Element Suchmaschinen zusätzliche Informationen über die Seite liefert. Es zeigt außerdem den Inhalt dieses `title`-Elements auf zwei weitere Arten an:

* in der Titelleiste, wenn die Seite geöffnet ist  
* im Browser-Tab der Seite, wenn Sie mit der Maus darüber fahren. Selbst wenn dieser Tab nicht aktiv ist, wird beim Darüberfahren der `title`-Text angezeigt.

Verschachteln Sie innerhalb des `head`-Elements ein `title`-Element mit dem Text `Colored Markers`.

# --hints--

Sie sollten ein öffnendes `title`-Tag haben.

```js
assert(code.match(/<title\s*>/i));
```

Sie sollten ein schließendes `title`-Tag haben.

```js
assert(code.match(/<\/title\s*>/i));
```

Ihr Projekt sollte den Titel `Colored Markers` haben.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

Denken Sie daran, dass Groß- und Kleinschreibung sowie Rechtschreibung für den Titeltext wichtig sind.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
