---
id: 61695197ac34f0407e339882
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Wie Sie in den vorherigen Projekten gesehen haben, sollten Webseiten mit einer `DOCTYPE html`-Deklaration beginnen, gefolgt von einem `html`-Element.

Fügen Sie eine `DOCTYPE html`-Deklaration am Anfang des Dokuments hinzu und danach ein `html`-Element. Geben Sie dem `html`-Element ein `lang`-Attribut mit `en` als Wert.

# --hints--

Ihr Code sollte die `DOCTYPE`-Referenz enthalten.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Sie sollten nach der `DOCTYPE`-Referenz ein Leerzeichen einfügen.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Sie sollten den Dokumenttyp als `html` definieren.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Sie sollten die `DOCTYPE`-Deklaration mit einem `>` nach dem Typ schließen.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Ihr `html`-Element sollte ein öffnendes Tag mit einem `lang`-Attribut von `en` haben.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ihr `html`-Element sollte ein schließendes Tag haben.

```js
assert(code.match(/<\/html\s*>/gi));
```

Ihre `DOCTYPE`-Deklaration sollte am Anfang Ihres HTML stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
