---
id: 61695ab9f6ffe951c16d03dd
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Verschachteln Sie ein `head`-Element innerhalb des `html`-Elements. Fügen Sie direkt nach dem `head`-Element ein `body`-Element hinzu.

# --hints--

Sie sollten ein öffnendes `head`-Tag haben.

```js
assert(code.match(/<head\s*>/i));
```

Sie sollten ein schließendes `head`-Tag haben.

```js
assert(code.match(/<\/head\s*>/i));
```

Sie sollten ein öffnendes `body`-Tag haben.

```js
assert(code.match(/<body\s*>/i));
```

Sie sollten ein schließendes `body`-Tag haben.

```js
assert(code.match(/<\/body\s*>/i));
```

Es sollte nur ein öffnendes `head`-Tag vorhanden sein.

```js
assert(code.match(/<head\s*>/ig).length === 1);
```

Es sollte nur ein schließendes `head`-Tag vorhanden sein.

```js
assert(code.match(/<\/head\s*>/ig).length === 1);
```

Es sollte nur ein öffnendes `body`-Tag vorhanden sein.

```js
assert(code.match(/<body\s*>/ig).length === 1);
```

Es sollte nur ein schließendes `body`-Tag vorhanden sein.

```js
assert(code.match(/<\/body\s*>/ig).length === 1);
```

Ihr `head`-Element sollte leer sein.

```js
assert(code.match(/<head\s*>\s*<\/head\s*>/i));
```

Ihr `body`-Element sollte leer sein.

```js
assert(code.match(/<body\s*>\s*<\/body\s*>/i));
```

Ihr `body`-Element sollte nach Ihrem `head`-Element platziert sein.

```js
assert(code.match(/<\/head\s*>\s*<body\s*>/i));
```

Ihr `head`-Element sollte innerhalb des `html`-Elements sein.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Ihr `body`-Element sollte innerhalb des `html`-Elements sein.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
