---
id: 61695d1fbc003856628bf561
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Das Attribut `charset` gibt die vom Dokument verwendete Zeichencodierung an. `utf-8` (Unicode Transformation Format – 8-bit) ist ein Zeichencodierungsstandard, der für die elektronische Kommunikation verwendet wird.

Verschachteln Sie innerhalb des `head`-Elements ein `meta`-Element mit dem Attribut `charset`, das auf `"utf-8"` gesetzt ist.

# --hints--

Sie sollten ein `meta`-Element haben.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 1);
```

Ihr `meta`-Element sollte ein Void-Element sein, es darf kein schließendes Tag `</meta>` haben.

```js
assert.notMatch(code, /<\/meta>/i);
```

Ihr `meta`-Element sollte ein Attribut `charset` mit dem Wert `"utf-8"` haben.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Ihr `meta`-Element sollte sich innerhalb des `head`-Elements befinden.

```js
const metaElementRegex = /<head\s*>(?:.|\r|\n)*?<meta(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, metaElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
