---
id: 616d3a67ccf800ad94ec89ae
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Jetzt fügen Sie einige Elemente hinzu, die Sie später als Farbmarker gestalten werden.

Fügen Sie zuerst innerhalb des `body`-Elements ein `div`-Element hinzu und setzen Sie dessen `class`-Attribut auf `container`. Stellen Sie sicher, dass das `div`-Element unterhalb des `h1`-Elements steht.

# --hints--

Sie sollten ein öffnendes `div`-Tag haben.

```js
assert(code.match(/<div.*>/i));
```

Sie sollten ein schließendes `div`-Tag haben.

```js
assert(code.match(/<\/div\s*>/i));
```

Ihr `div`-Element sollte sich innerhalb Ihres `body`-Elements befinden.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

Ihr `div`-Element sollte nach dem `h1`-Element stehen.

```js
assert.exists(document.querySelector('h1 + div'));
```

Ihr `div`-Element sollte ein `class`-Attribut mit dem Wert `container` haben.

```js
assert(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
--fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
--fcc-editable-region--
</html>
```

```css
h1 {
  text-align: center;
}
```
