---
id: 61764c602bee6974e7790f35
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Um den Markierungen unterschiedliche Farben zu geben, müssen Sie jeder eine eindeutige Klasse hinzufügen. Mehrere Klassen können einem Element hinzugefügt werden, indem Sie sie im `class`-Attribut auflisten und durch ein Leerzeichen trennen. Zum Beispiel fügt das Folgende sowohl die `animal`- als auch die `dog`-Klasse zu einem `div`-Element hinzu.

```html
<div class="animal dog">
```

Wenn Sie mehreren Klassen zu einem HTML-Element hinzufügen, können die Stile der zuerst aufgelisteten Klassen von späteren Klassen überschrieben werden.

Beginnen Sie damit, die Klasse `one` zum ersten Marker-`div`-Element hinzuzufügen.


# --hints--

Sie sollten die Klasse `one` zum ersten Marker-`div`-Element hinzufügen.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

Ihr erster Marker-`div` sollte die Klassen `marker` und `one` haben.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
