---
id: 617b9ad735109e217284e095
title: Schritt 27
challengeType: 0
dashedName: step-27
---

# --description--

Jetzt, da Sie die primären RGB-Farben haben, ist es Zeit, sie zu kombinieren.

Setzen Sie für die `rgb`-Funktion in der `.container`-Regel die Rot-, Grün- und Blauwerte auf das Maximum von `255`.

# --hints--

Ihre `.container`-CSS-Regel sollte eine `background-color`-Eigenschaft mit dem Wert `rgb(255, 255, 255)` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.container {
  background-color: rgb(0, 0, 0);
  padding: 10px 0;
}
--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 255, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
