---
id: 617b954d9f4f6217a749380e
title: Schritt 25
challengeType: 0
dashedName: step-25
---

# --description--

Fügen Sie nun etwas mehr vertikalen Abstand zwischen Ihren Markierungen und dem Rand des `container`-Elements hinzu, in dem sie sich befinden.

Verwenden Sie in der `.container`-CSS-Regel die Kurzschreibweise der `padding`-Eigenschaft, um `10px` Innenabstand oben und unten hinzuzufügen, und setzen Sie den Innenabstand links und rechts auf `0`. Dies funktioniert ähnlich wie die Kurzschreibweise der `margin`-Eigenschaft, die Sie zuvor verwendet haben.

# --hints--

Sie sollten die `background-color`-Eigenschaft und ihren Wert in der `.container`-CSS-Regel nicht entfernen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

Ihre `.container`-CSS-Regel sollte eine `padding`-Eigenschaft mit dem Wert `10px 0` enthalten.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.padding === '10px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.container {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 127, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
