---
id: 619b7c3c83de403126b69c1e
title: Schritt 47
challengeType: 0
dashedName: step-47
---

# --description--

Sie kennen vielleicht bereits Dezimalwerte oder Basis-10-Werte, die von 0 bis 9 reichen. Hexadezimalwerte oder Basis-16-Werte reichen von 0 bis 9 und dann von A bis F:

```js
0, 1, 2, 3, 4, 5, 6, 7, 8, 9, A, B, C, D, E, F
```

Bei Hex-Farben steht `00` für 0 % dieser Farbe und `FF` für 100 %. Also bedeutet `#00FF00` 0 % Rot, 100 % Grün und 0 % Blau und entspricht `rgb(0, 255, 0)`.

Verringern Sie die Intensität von Grün, indem Sie den Grünwert der Hex-Farbe auf `7F` setzen.

# --hints--

Ihre `.green`-CSS-Regel sollte eine `background-color`-Eigenschaft mit dem Wert `#007F00` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.green {
  background-color: #00FF00;
}
--fcc-editable-region--

.blue {
  background-color: rgb(0, 0, 0);
}

```
