---
id: 61b0936d9e7fc020b43b1bb8
title: Schritt 78
challengeType: 0
dashedName: step-78
---

# --description--

Es sieht so aus, als wäre Ihr Ärmel verschwunden, aber keine Sorge – er ist noch da. Was passiert ist: Ihre neue Kappe `div` nimmt die gesamte Breite des Markers ein und drückt den Ärmel in die nächste Zeile.

Das liegt daran, dass die Standard-`display`-Eigenschaft für `div`-Elemente `block` ist. Wenn also zwei `block`-Elemente nebeneinanderstehen, stapeln sie sich wie echte Blöcke. Zum Beispiel sind Ihre Marker-Elemente alle übereinander gestapelt.

Um zwei `div`-Elemente in einer Zeile anzuordnen, setzen Sie deren `display`-Eigenschaften auf `inline-block`.

Erstellen Sie eine neue Regel, die sowohl die `cap`- als auch die `sleeve`-Klasse anspricht, und setzen Sie `display` auf `inline-block`.

# --hints--

Sie sollten einen Klassenselektor verwenden, um sowohl die `cap`- als auch die `sleeve`-Klasse anzusprechen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve') || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap'));
```

Ihre CSS-Regel sollte eine `display`-Eigenschaft mit dem Wert `inline-block` enthalten.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve')?.display === 'inline-block' || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap')?.display === 'inline-block');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
}

--fcc-editable-region--

--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
