---
id: 61b0a44a6b865738ba49b9fb
title: Schritt 86
challengeType: 0
dashedName: step-86
---

# --description--

Das Letzte, was Sie tun, ist, jedem Marker einen leichten Schatten hinzuzufügen, damit sie noch realistischer aussehen.

Die `box-shadow`-Eigenschaft ermöglicht es Ihnen, einen oder mehrere Schatten um ein Element herum anzuwenden. Hier ist die grundlegende Syntax:

```css
box-shadow: offsetX offsetY color;
```

So funktionieren die Werte `offsetX` und `offsetY`:

* Sowohl `offsetX` als auch `offsetY` akzeptieren Zahlenwerte in `px` und anderen CSS-Einheiten
* Ein positiver `offsetX`-Wert verschiebt den Schatten nach rechts, ein negativer Wert nach links
* Ein positiver `offsetY`-Wert verschiebt den Schatten nach unten, ein negativer Wert nach oben
* Wenn Sie für `0` und/oder `offsetX` den Wert null (`offsetY`) verwenden möchten, brauchen Sie keine Einheit anzugeben. Jeder Browser versteht, dass null keine Änderung bedeutet.

Die Höhe und Breite des Schattens wird durch die Höhe und Breite des Elements bestimmt, auf das er angewendet wird. Sie können auch einen optionalen `spreadRadius`-Wert verwenden, um die Ausbreitung des Schattens zu vergrößern. Mehr dazu später.

Beginnen Sie damit, dem roten Marker einen einfachen Schatten hinzuzufügen.

Fügen Sie in der `.red`-CSS-Regel die `box-shadow`-Eigenschaft mit den Werten `5px` für `offsetX`, `5px` für `offsetY` und `red` für `color` hinzu.

# --hints--

Ihre `.red`-CSS-Regel sollte eine `box-shadow`-Kurzschreibweise-Eigenschaft mit dem Wert `5px 5px red` enthalten.

```js
assert(new __helpers.CSSHelp(document).getStyle('.red')?.boxShadow === 'red 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

--fcc-editable-region--
.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}
--fcc-editable-region--

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
