---
id: a39963a4c10bc8b4d4f06d7e
title: 値を削除する関数を実装する
challengeType: 26
dashedName: implement-a-value-remover-function
---

# --description--

このラボでは、最初の引数に配列を受け取り、その後に削除する値を表す1つ以上の追加の引数を受け取る関数を作成します。

**目的**：以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー：**

1. `destroyer` 関数は配列と1つ以上の追加の引数を受け取る必要があります。
1. `destroyer` 関数は、最初の引数の中から後続の引数のいずれかに一致する要素を除外した新しい配列を返す必要があります。
1. 関数は不定数の引数を受け取る必要があります。

# --hints--

`destroyer` 関数を用意してください。

```js
assert.isFunction(destroyer);
```

`destroyer([1, 2, 3, 1, 2, 3], 2, 3)` は `[1, 1]` を返すべきです。

```js
assert.deepEqual(destroyer([1, 2, 3, 1, 2, 3], 2, 3), [1, 1]);
```

`destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3)` は `[1, 5, 1]` を返すべきです。

```js
assert.deepEqual(destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3), [1, 5, 1]);
```

`destroyer([3, 5, 1, 2, 2], 2, 3, 5)` は `[1]` を返すべきです。

```js
assert.deepEqual(destroyer([3, 5, 1, 2, 2], 2, 3, 5), [1]);
```

`destroyer([2, 3, 2, 3], 2, 3)` は `[]` を返すべきです。

```js
assert.deepEqual(destroyer([2, 3, 2, 3], 2, 3), []);
```

`destroyer(["tree", "hamburger", 53], "tree", 53)` は `["hamburger"]` を返すべきです。

```js
assert.deepEqual(destroyer(["tree", "hamburger", 53], "tree", 53), ["hamburger"]);
```

`destroyer( ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"], "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan" )` は `[12, 92, 65]` を返すべきです。

```js
assert.deepEqual(
  destroyer(
    ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"],
    "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan"
  ),
  [12, 92, 65]
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function destroyer(arr, ...valsToRemove) {
  return arr.filter(item => !valsToRemove.includes(item));
}
```
