---
id: 65f569725359e10d345bc52a
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

Sie werden ein Programm erstellen, das die Flugbahn eines Projektils berechnet und zeichnet, basierend auf dem Winkel, der Geschwindigkeit und der Höhe des Wurfs.

Beginnen Sie damit, `math` zu importieren. Sie werden es in diesem Projekt häufig verwenden, da es nützliche Methoden wie `math.radians`, `math.cos`, `math.sin` und andere enthält.

Erstellen Sie außerdem diese Variablen, um den Wert der Erdbeschleunigung und einige spezielle Symbole zu speichern, die später nützlich sein werden (verwenden Sie zum Einfügen Kopieren und Einfügen).

```py
GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
```

# --hints--

Verwenden Sie `import math`, um das `math`-Modul zu importieren.

```js
({test: () => assert(runPython(`_Node(_code).has_import('import math')`))})
```

Sie sollten eine `GRAVITATIONAL_ACCELERATION`-Variable mit dem Wert `9.81` haben.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("GRAVITATIONAL_ACCELERATION"), "missing variable"
assert _Node(_code).find_variable("GRAVITATIONAL_ACCELERATION").is_equivalent('GRAVITATIONAL_ACCELERATION = 9.81'), "variable has wrong value"
`)})
```

Sie sollten eine `PROJECTILE`-Variable mit dem Wert `"∙"` haben.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("PROJECTILE"), "missing variable"
assert _Node(_code).find_variable("PROJECTILE").is_equivalent('PROJECTILE = "∙"'), "variable has wrong value"
`)})
```

Sie sollten eine `x_axis_tick`-Variable mit dem Wert `"T"` haben.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("x_axis_tick"), "missing variable"
assert _Node(_code).find_variable("x_axis_tick").is_equivalent('x_axis_tick = "T"'), "variable has wrong value"
`)})
```

Sie sollten eine `y_axis_tick`-Variable mit dem Wert `"⊣"` haben.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("y_axis_tick"), "missing variable"
assert _Node(_code).find_variable("y_axis_tick").is_equivalent('y_axis_tick = "⊣"'), "variable has wrong value"
`)})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
