---
id: 65fd4bd84561a14d3e720061
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Die Klassenvariable `__slots__` hat eine besondere Verwendung in Python-Klassen. Wenn Sie `__slots__` deklarieren und ihr eine Folge von Strings zuweisen, beschränkt das die Erstellung von Attributen auf diejenigen, die in dieser Folge enthalten sind. Außerdem verhindert es die Erstellung des speziellen Attributs `__dict__` und ermöglicht einen effizienteren Zugriff auf Attribute.

Sie sollten die Variable `__slots__` innerhalb der Klasse verwenden, um zu definieren, welche Attribute die Klasse hat: Weisen Sie `__slots__` ein Tuple zu, das 3 Strings enthält, von denen jeder einem der in `__init__` definierten Attributnamen entspricht.

# --hints--

Sie sollten `__slots__` ein Tuple zuweisen, das `'__height'`, `'__speed'` und `'__angle'` enthält.

```js
({test: () => assert(runPython(`
from itertools import permutations
slots = _Node(_code).find_class("Projectile").find_body().find_variable("__slots__")
perms = permutations(("__height", "__speed", "__angle"))
any(slots.is_equivalent(f'__slots__ = {perm}') for perm in perms)
`))})

```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--
class Projectile:

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
--fcc-editable-region--
```
