---
id: 65fd4b3bedee044b4b957d33
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Erstellen Sie eine `Projectile`-Klasse mit einer `__init__`-Methode, die die Klasse mit drei Argumenten initialisiert: der Anfangsgeschwindigkeit, der Anfangshöhe und dem Anfangswinkel des Wurfs des Projektils, in dieser Reihenfolge.

Weisen Sie innerhalb der `__init__`-Methode diese Argumente drei privaten Attributen zu: `__speed`, `__height` und `__angle`. Der angegebene Winkel wird in Grad sein; er sollte jedoch intern in Radiant gespeichert werden. Verwenden Sie dazu die `math.radians()`-Funktion, um den Winkel beim Zuweisen an `__angle` von Grad in Radiant umzuwandeln.

Die Verwendung von zwei Unterstrichen vor einem Attributnamen löst in Python Name Mangling aus. Das bedeutet, dass die Attribute von außen nicht direkt mit ihren angegebenen Namen zugreifbar sind und bei Bedarf mit den gemangelten Namen wie `ball._Projectile__height` aufgerufen werden müssen, was anzeigt, dass diese nur für den internen Gebrauch vorgesehen sind.

# --hints--

Sie sollten `class Projectile:` verwenden, um die Klasse zu erstellen.

```js
({test: () => assert(runPython(`_Node(_code).has_class("Projectile")`))})

```

Sie sollten eine `__init__`-Methode innerhalb der `Projectile`-Klasse haben.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_body().has_function("__init__")`))})
```

Eine neue `Projectile`-Klasseninstanz sollte die Werte der drei Attribute korrekt initialisiert haben.

```js
({test: () => assert(runPython(`
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
p = Projectile(20, 21, 22)
p._Projectile__height == 21 and p._Projectile__speed == 20 and p._Projectile__angle == math.radians(22)
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--

--fcc-editable-region--
```
