---
id: 660fcf3e1b9bb056b2edb567
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

An diesem Punkt sind Sie bereit, die String-Darstellung zu erstellen.  
Beginnen Sie damit, eine Instanz der `Projectile`-Klasse zu erstellen. Unterhalb der Klassendefinition erstellen Sie eine `ball`-Variable und weisen ihr einen Aufruf von `Projectile` mit `10, 3, 45` als Argumente zu.

Definieren Sie innerhalb der Klasse eine `__str__`-Methode, um Ihrer Instanz eine korrekte String-Darstellung zu geben, und lassen Sie sie einen String mit folgendem Format zurückgeben:

```py

Projectile details:
speed: 10 m/s
height: 3 m
angle: 45°
displacement: 12.6 m

```

Der String soll mit einem Zeilenumbruch beginnen und enden, der Winkel muss als ganze Zahl in Grad geschrieben werden, und die Verschiebung soll mit einer Dezimalziffer ausgegeben werden.

Sie finden nützliche `math.degrees`, um den Winkel von Radiant in Grad umzuwandeln.

Wenn Sie bereit sind, können Sie `print(ball)` verwenden, um Ihre Funktion zu testen.

# --hints--

Sie sollten eine `ball`-Variable haben.

```js
({test: () => assert(runPython(`_Node(_code).has_variable('ball')`))})
```

Die `ball`-Variable sollte den Wert `Projectile(10, 3, 45)` haben.

```js
({test: () => assert(runPython(`_Node(_code).find_variable('ball').is_equivalent('ball = Projectile(10, 3, 45)')`))})
```

Die String-Darstellung für `Projectile(45, 45, 45)` sollte korrekt sein.

```js
({test: () => assert(runPython(
`
ball = Projectile(45, 45, 45)
str(ball) == """
Projectile details:
speed: 45 m/s
height: 45 m
angle: 45°
displacement: 244.4 m
"""`
))})
```

Die String-Darstellung sollte auch für andere Instanzen korrekt sein.

```js
({test: () => assert(runPython(`p = Projectile(10, 10, 10)
str(
    p
) == """
Projectile details:
speed: 10 m/s
height: 10 m
angle: 10°
displacement: 15.9 m
"""`))})
```

Sie sollten einen `print(ball)`-Aufruf haben.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls('print')
    assert any(c.is_equivalent('print(ball)') for c in calls)
    `)
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)


    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
--fcc-editable-region--

--fcc-editable-region--
        
```
