---
id: 6616ec0bf9ee8548ce6b0f08
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Jetzt ist es Zeit, an der Berechnung der Koordinaten der Flugbahn zu arbeiten! Erstellen Sie eine Methode der `Projectile`-Klasse mit dem Namen `__calculate_y_coordinate`, die neben `self` einen `x`-Parameter haben soll.

\\[ y = y_0 + x \tan\theta - \frac{g x^2}{2 v_0^2 \cos^2\theta} \\]

Oben sehen Sie die Formel zur Berechnung der vertikalen Position $y$ für eine beliebige horizontale Position $x$, wobei der Startwinkel $\theta$, die Geschwindigkeit $v_0$ und die Höhe $y_0$ gegeben sind.

Sie müssen `math.tan()` und `math.cos()` verwenden und daran denken, dass `x ** y` die Schreibweise für $x^y$ ist und dass der Wert von $g$ in der Variablen `GRAVITATIONAL_ACCELERATION` gespeichert ist.

Implementieren Sie die Methode so, dass sie die $y$-Koordinate zurückgibt.

Wenn Sie die Methode testen möchten, können Sie sie von außerhalb der Klasse aufrufen, indem Sie `ball._Projectile__calculate_y_coordinate()` mit einer Zahl als Argument aufrufen, nachdem `ball` deklariert wurde.

# --hints--

Ihre `Projectile`-Klasse sollte eine Methode mit dem Namen `__calculate_y_coordinate` haben.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').has_function('__calculate_y_coordinate')`))}})
```

Die `__calculate_y_coordinate`-Methode sollte `self, x` als Argumente haben.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').find_function('__calculate_y_coordinate').has_args('self, x')`))}})
```

Die Methode sollte den korrekten Wert zurückgeben. Für `Projectile(12, 12, 12)` sollten bei den 4 $x$-Koordinaten 0, 1, 2 und 3 die y-Koordinaten ungefähr 12, 12,18, 12,28, 12,32 sein.

```js
({test: () => {
    assert(runPython(`
    a = Projectile(12, 12, 12)
    all([round(a._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(0, 12.0), (1, 12.18), (2, 12.28), (3, 12.32)]])
    `))
}})
```

Die Methode sollte den korrekten Wert zurückgeben. Für `Projectile(45, 12, 22)` sollten bei den $x$-Koordinaten 6, 9, 10 und 12 die y-Koordinaten ungefähr 14,32, 15,41, 15,76, 16,44 sein.

```js
({test: () => {
    assert(runPython(`
    b = Projectile(45, 12, 22)
    all([round(b._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(6, 14.32), (9, 15.41), (10, 15.76), (12, 16.44)]])
    `))
}})
```

Die `__calculate_y_coordinate`-Methode sollte das Ergebnis nicht runden.

```js
({test: () => assert.isFalse(runPython(`Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5) == round(Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5), 2)`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    --fcc-editable-region--
    
    --fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
   
```
