---
id: 662fa453ec8033322b3ab2e3
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Es ist Zeit, über Kapselung und Getter zu sprechen. Sie haben die drei Instanzattribute mit einem führenden doppelten Unterstrich als privat definiert. Beachten Sie, dass diese Attribute nach Konvention als privat bezeichnet werden: Obwohl sie von außen weiterhin aufgerufen werden können, ist es üblich, dies nicht zu tun.

Getter werden verwendet, um die Werte von außen zu erhalten. Um einen Getter zu definieren, definieren Sie eine Methode, die den Wert des gewünschten Attributs zurückgibt, und versehen Sie sie mit einem `@property`-Dekorator:

```py
class Nest:
    ...
    @property
    def number_of_eggs(self):
        return self.__number_of_eggs
```

 Der Dekorator wandelt die Methode in eine Eigenschaft um, was bedeutet, dass die Methode nicht wie eine reguläre Methode aufgerufen wird, sondern wie eine Eigenschaft verwendet wird:

```py
n = Nest()
print(n.number_of_eggs)
```

Im obigen Beispiel wird auf das private Attribut `__number of eggs` über die `number_of_eggs`-Eigenschaft von `n` zugegriffen.

Erstellen Sie einen Getter namens `speed` und lassen Sie ihn den Wert des privaten Attributs `__speed` zurückgeben.

# --hints--

Sie sollten eine `speed`-Methode definieren.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').has_function('speed')
`))})
```

Ihre `speed`-Methode sollte einen einzelnen Parameter, `self`, haben.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_args('self')
`))})
```

Ihre `speed`-Methode sollte einen `@property`-Dekorator haben.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_decorators('property')
`))})
```

Ihre `speed`-Methode sollte den Wert des privaten Attributs `__speed` zurückgeben.

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.speed == 22
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

--fcc-editable-region--
    
--fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()
   
```
