---
id: 6633b26c4f2d603072ef0818
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Es ist gute Praxis, der Klasse eine Darstellung zu geben, indem Sie die spezielle Methode `__repr__` verwenden. Während die Methode `__str__` eine lesbare String-Darstellung zurückgibt, die benutzerfreundlich sein soll, ist `__repr__` für Programmierende gedacht. Oft liefert `__repr__` einen String, der verwendet werden kann, um das Objekt neu zu erstellen.

Schreiben Sie die Methode `__repr__`, die den String zurückgibt, der benötigt wird, um das Objekt zu instanziieren.

Sie können dies testen, indem Sie die Funktion `repr` aufrufen und ihr eine Instanz der Klasse übergeben.

# --hints--

Die Klasse `Projectile` sollte eine Methode `__repr__` haben.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_class('Projectile').has_function('__repr__')`
    ))
})
```

Die Instanz `Projectile(12, 12, 12)` sollte eine Darstellung gleich `'Projectile(12, 12, 12)'` haben.

```js
({
    test: () => assert(runPython(
        `repr(Projectile(12, 12, 12)) == 'Projectile(12, 12, 12)' or \\
        repr(Projectile(12, 12, 12)) == 'Projectile(12, 12, 12.0)'`
    ))
})
```

Die Instanz `Projectile(21, 22, 45)` sollte eine Darstellung gleich `'Projectile(21, 22, 45)'` haben.

```js
({
    test: () => assert(runPython(
        `repr(Projectile(21, 22, 45)) == 'Projectile(21, 22, 45)' or \\
        repr(Projectile(21, 22, 45)) == 'Projectile(21, 22, 45.0)'`
    ))
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.speed} m/s
height: {self.height} m
angle: {self.angle}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

    @property
    def height(self):
        return self.__height

    @property
    def angle(self):
        return round(math.degrees(self.__angle))

    @property
    def speed(self):
        return self.__speed
    
    @height.setter
    def height(self, n):
        self.__height = n

    @angle.setter
    def angle(self, n):
        self.__angle = math.radians(n)

    @speed.setter
    def speed(self, s):
       self.__speed = s

--fcc-editable-region--
    
--fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()

   
```
