---
id: 69a5f35669099ed52f8563b1
title: Jenga programu ya kujaza tena pantri kwa akili
challengeType: 26
dashedName: lab-smart-pantry-restocker
---

# --description--

Katika maabara hii, utajenga programu ndogo ya usimamizi wa pantri ukitumia dhana za msingi za JavaScript kama safu za data, vitu, mizunguko, na masharti.

Utahakikisha kupokea mzigo wa vitu vya pantri, kuamua nini cha kufanya na kila kipengee, na kupanga matokeo kwa ajili ya kuhifadhi.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara hii.

Safu ya `rawData` ina mfuatano wa herufi ulio na vipengele vilivyounganishwa kwa alama ya bomba kwa muundo `sku|name|qty|expires|zone`, ambapo `zone` ni hiari.

**Hadithi za watumizi:**

1. Unapaswa kutekeleza kitendakazi `parseShipment(rawData)` kinachopokea safu ya mfuatano wa herufi na kurudisha safu ya vitu vyenye sifa za `{ sku, name, qty, expires, zone }`.

   - Thamani za `sku` zinazojirudia katika mzigo zinapaswa kupuuzwa.
   - Wakati sehemu ya `zone` haijatolewa, inapaswa kuwekwa kuwa `"general"`.
   - Thamani ya `qty` inapaswa kubadilishwa kuwa nambari.

2. Unapaswa kutekeleza kitendakazi `planRestock(pantry, shipment)` kinacholinganishwa pantri ya sasa na mzigo unaokuja na kurudisha safu ya vitendo kwa aina ya `{ type, item }`, ambapo `type` ni mojawapo ya `"restock"`, `"discard"`, au `"donate"`, na `item` ni kipengee kilichochanganuliwa cha mzigo.

   Kigezo `pantry` ni safu ya vitu vyenye muundo sawa na kipengee cha mzigo kilichochanganuliwa (`{ sku, name, qty, expires, zone }`).

- Ikiwa kipengee cha mzigo kina `qty` ya `0` au chini, aina ya kitendo `type` inapaswa kuwa `"discard"`, bila kujali kama kipengee kiko pantri.
   - Vinginevyo, ikiwa `sku` ya kipengee cha mzigo tayari ipo pantri, aina ya kitendo `type` inapaswa kuwa `"restock"`.
   - Vinginevyo (ikiwa `sku` ya kipengee cha mzigo haipo pantri), aina ya kitendo `type` inapaswa kuwa `"donate"`.

3. Unapaswa kutekeleza kitendakazi `groupByZone(actions)` kinachokusanya vitendo katika maeneo ya kuhifadhi kulingana na sifa ya `zone` ya kila kipengee. Kitendakazi kinapaswa kurudisha kitu cha aina ya kitu ambacho kila ufunguo ni jina la eneo na thamani ni safu ya vitendo vinavyohusiana na eneo hilo. Kwa mfano, ikiwa vitendo vina vitu vyenye maeneo `"fridge"` na `"pantry"`, matokeo yanapaswa kuwa `{ fridge: [...], pantry: [...] }`.

4. Unapaswa kutekeleza kitendakazi `clonePantry(pantry)` kinachorudisha nakala ya kina ya pantri ili mabadiliko ya kupanga yasibadilishe orodha ya awali. Nakala ya kina inamaanisha kuunda safu mpya yenye vitu vipya, hivyo kubadilisha nakala hakubadilishi pantri ya awali.

5. Unapaswa kutumia vitendakazi vyote pamoja kuchakata mzigo na kuandika kwenye konsoli kitu cha mwisho cha matokeo yaliyokusanywa.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kufafanua kitendakazi kinachoitwa `parseShipment` kinachopokea safu moja ya mfuatano wa herufi inayoitwa `rawData` kama kigezo.

``` js
assert.isFunction(parseShipment);
assert.lengthOf(parseShipment, 1);
```

Kitendakazi chako cha `parseShipment` kinapaswa kubadilisha mfuatano wa herufi wa mzigo katika safu kuwa vitu vyenye sifa: `sku`, `name`, `qty`, `expires`, na `zone`.

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);

assert.isArray(result);
assert.isObject(result[0]);
assert.deepEqual(result[0], {
  sku: "A10",
  name: "Tomatoes",
  qty: 5,
  expires: "2027-01-01",
  zone: "fridge"
});
```

SKU zinazojirudia katika mzigo zinapaswa kupuuzwa.

``` js
const rawData = [
  "C32|Eggs|3|2027-01-01",
  "C32|Eggs|3|2027-01-01"
];

const result = parseShipment(rawData);

assert.lengthOf(result, 1);
assert.strictEqual(result[0].sku, "C32");
```

Wakati `zone` haipo katika mfuatano wa mzigo, unapaswa kuziweka kuwa `general`.

``` js
const rawData = ["B21|Bananas|10|2027-01-01"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].zone, "general");
```

Thamani ya `qty` inapaswa kubadilishwa kuwa nambari.

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].qty, 5);
assert.isNumber(result[0].qty);
```

Unapaswa kufafanua kitendakazi kinachoitwa `planRestock` kinachopokea vigezo viwili: `pantry` na `shipment`.

``` js
assert.isFunction(planRestock);
assert.lengthOf(planRestock, 2);
```

Kitendakazi chako cha `planRestock` kinapaswa kurudisha safu ya vitendo vyenye sifa za `type` na `item`.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 6, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.isArray(result);
assert.isObject(result[0]);
assert.property(result[0], "type");
assert.property(result[0], "item");
```

Vitu vyenye kiasi cha `0` au chini katika mzigo vinapaswa kuunda kitendo cha aina ya `discard`, bila kujali kama kipengee kiko pantri.

``` js
const pantry = [{ sku: "D43", name: "Pineapples", qty: 2, expires: "2027-01-01", zone: "general" }];
const shipment = [{ sku: "D43", name: "Pineapples", qty: 0, expires: "2027-01-01", zone: "general" },
                  { sku: "E54", name: "Peppers", qty: -1, expires: "2027-01-01", zone: "fridge" }
                ];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "discard");
assert.strictEqual(result[1].type, "discard");
```

Wakati kipengee cha mzigo tayari kiko pantri (SKU sawa na `sku`), kitendo kinapaswa kuwa `restock`.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "restock");
```

Kama kipengee cha mzigo hakipo pantri, kitendo kinapaswa kuwa `"donate"`.

``` js
const pantry = [];
const shipment = [{ sku: "C32", name: "Eggs", qty: 3, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "donate");
```

Unapaswa kufafanua kitendakazi kinachoitwa `groupByZone` kinachopokea kigezo kimoja kinachoitwa `actions`.

``` js
assert.isFunction(groupByZone);
assert.lengthOf(groupByZone, 1);
```

Kitendakazi chako cha `groupByZone` kinapaswa kurudisha vitendo vilivyokusanywa kulingana na sifa ya `zone` ya kila kitu.

``` js

const actions = [
  { type: "restock", item: { sku: "A1", zone: "fridge" } },
  { type: "restock", item: { sku: "B1", zone: "pantry" } }, 
];

const result = groupByZone(actions);

assert.isObject(result);
assert.isArray(result.fridge);
assert.isArray(result.pantry);
```

Kitendakazi chako cha `groupByZone` kinapaswa kukusanya vitendo vyenye maudhui na idadi sahihi.

``` js
const actions = [
  { type: "restock", item: { sku: "A1", zone: "fridge" } },
  { type: "donate",  item: { sku: "B1", zone: "fridge" } },
  { type: "discard", item: { sku: "C1", zone: "pantry" } },
];

const result = groupByZone(actions);

assert.lengthOf(result.fridge, 2);
assert.lengthOf(result.pantry, 1);
assert.strictEqual(result.fridge[0].item.sku, "A1");
assert.strictEqual(result.pantry[0].type, "discard");
```

Unapaswa kufafanua kitendakazi kinachoitwa `clonePantry` kinachopokea kigezo kimoja kinachoitwa `pantry`.

``` js
assert.isFunction(clonePantry);
assert.lengthOf(clonePantry, 1);
```

`clonePantry` inapaswa kurudisha safu mpya badala ya safu ya pantri ya awali.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.isArray(pantryClone);
assert.notStrictEqual(pantryClone, pantry);
```

Vitu ndani ya pantri iliyokopiwa pia vinapaswa kuwa vitu vipya (nakala ya kina).

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.notStrictEqual(pantryClone[0], pantry[0]);
assert.deepEqual(pantryClone[0], pantry[0]);

pantryClone[0].qty = 100;
assert.strictEqual(pantry[0].qty, 4);
```

Vitendakazi vinapaswa kufanya kazi pamoja kuchakata mzigo na kukusanya vitendo vinavyotokea.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];

const shipment = parseShipment(rawData);
const pantryClone = clonePantry(pantry);
const actions = planRestock(pantryClone, shipment);
const grouped = groupByZone(actions);

assert.isObject(grouped);
assert.strictEqual(grouped.fridge[0].item.qty, 5);
```

Unapaswa kuandika vitendo vilivyokusanywa kulingana na maeneo ya kuhifadhi.

``` js
const groupedLog = getLogs().find(log => {
  if (!log || typeof log !== "object" || Array.isArray(log)) return false;

  return Object.values(log).some(zoneActions => {
    if (!Array.isArray(zoneActions)) return false;

    return zoneActions.some(action =>
      action &&
      ["restock", "discard", "donate"].includes(action.type) &&
      action.item &&
      typeof action.item.sku === "string"
    );
  });
});

assert.isOk(groupedLog, "Log the grouped actions object to the console.");
```

# --seed--

## --seed-contents--

```js
const pantry = [
  { sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" },
  { sku: "D43", name: "Pineapples", qty: 2, expires: "2020-01-01", zone: "general" }
];

const rawData = [
  "A10|Tomatoes|5|2027-01-01",
  "B21|Bananas|10|2027-01-01",
  "C32|Eggs|3|2027-01-01|fridge",
  "C32|Eggs|3|2027-01-01",
  "D43|Pineapples|0|2027-01-01",
  "E54|Peppers|-1|2027-01-01|fridge"
];
```

# --solutions--

```js

function parseShipment(rawData) {

  // Early return if rawData is not an Array
  if (!Array.isArray(rawData)) return [];

  const result = [];
  const seenSKUs = [];

  for (let line of rawData) {

    const parts = line.split("|");

    const sku = parts[0].trim();

    if (seenSKUs.includes(sku)) continue; // Avoid duplicates

    const name = parts[1].trim();
    const qty = Number(parts[2].trim()); // Parse qty into a number
    const expires = parts[3].trim();
    const zone = ( parts[4] ) ? parts[4].trim() : "general";

    result.push({
      sku: sku,
      name: name,
      qty: qty,
      expires: expires,
      zone: zone
    });

    seenSKUs.push(sku);

  }

  return result;
}

function planRestock(pantry, shipment) {
  const actions = [];

  for (let item of shipment) {

    var exists = false;

    for (let stock of pantry)
    {
      if (stock.sku === item.sku)
      {
        exists = true;
      }
    }

    if (item.qty > 0) {
      if (exists) {
        actions.push({ type: "restock", item: item });
      } else{
        actions.push({ type: "donate", item: item });
      }
    } else {
      actions.push({ type: "discard", item: item });
    }
  }

  return actions;
}

function groupByZone(actions) {

  const byZone = {};

  for (const action of actions) {

    const zone = action.item.zone;

    if (!byZone[zone]) {
      byZone[zone] = [];
    }

    byZone[zone].push(action);
  }

  return byZone;

}

function clonePantry(pantry) {
  const pantryClone = [];

  for (let item of pantry) {
    pantryClone.push({
      sku: item.sku,
      name: item.name,
      qty: item.qty,
      expires: item.expires,
      zone: item.zone
    });
  }

  return pantryClone;
}

const pantry = [
  { sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" },
  { sku: "D43", name: "Pineapples", qty: 2, expires: "2020-01-01", zone: "general" }
];

const rawData = [
  "A10|Tomatoes|5|2027-01-01", // Restock existing item
  "B21|Bananas|10|2027-01-01", // Donate new item without zone
  "C32|Eggs|3|2027-01-01|fridge", // Donate to a defined zone
  "C32|Eggs|3|2027-01-01", // Duplicated SKU in shipment
  "D43|Pineapples|0|2027-01-01", // Discard with quantity of 0
  "E54|Peppers|-1|2027-01-01|fridge" // Discard even if it's not in pantry
];

const shipment = parseShipment(rawData);
const pantryCopy = clonePantry(pantry);
const actions = planRestock(pantryCopy, shipment);
const grouped = groupByZone(actions);

// Show items expanded in zones
console.log(grouped);

```
