---
id: 5dc1798ff86c76b9248c6eb3
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Debajo del elemento `h1`, añade un elemento `h2` con este texto:

`Cat Photos`

# --hints--

Tu elemento `h1` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Tu elemento `h1` debe tener una etiqueta de cierre. Las etiquetas de cierre tienen esta sintaxis: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Debes tener solo un elemento `h1`. Elimina el extra.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

El texto de tu elemento `h1` debe ser 'CatPhotoApp'. O has omitido el texto o tienes un error tipográfico.

```js
assert.equal(document.querySelector('h1')?.innerText.toLowerCase(), 'catphotoapp');
```

Tus etiquetas `h2` deben estar en minúsculas. Por convención, todas las etiquetas HTML se escriben en minúsculas.

```js
assert.notMatch(code, /<\/?H2>/);
```

Tu elemento `h2` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector('h2'));
```

Tu elemento `h2` debe tener una etiqueta de cierre. Las etiquetas de cierre tiene una `/` después del carácter `<`.

```js
assert.match(code, /<\/h2\>/);
```

El texto de tu elemento `h2` debe de ser `Cat Photos`. Solo coloca el texto `Cat Photos` entre las etiquetas de apertura y cierre del `h2`.

```js
assert.equal(document.querySelector('h2')?.innerText.toLowerCase(), 'cat photos');
```

Tu elemento `h2` debe estar debajo del elemento `h1`. El elemento `h1` tiene mayor importancia y debe estar encima del elemento `h2`.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H1'), collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>
    
--fcc-editable-region--
  </body>
</html>
```
