---
id: 5dc17d3bf86c76b9248c6eb4
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Crea un elemento `p` debajo de tu elemento `h2` y dale el siguiente texto:

`Everyone loves cute cats online!`

# --hints--

Tu elemento `p` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen la siguiente sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector('p'));
```

Tus etiquetas `p` deben estar en minúsculas. Por convención, todas las etiquetas HTML se escriben en minúsculas.

```js
assert.notMatch(code, /<\/?P>/);
```

Tu elemento `p` debe tener una etiqueta de cierre. Las etiquetas de cierre tiene una `/` después del carácter `<`.

```js
assert.match(code, /<\/p\>/);
```

El texto de tu elemento `p` debería ser `¡A todos les encantan los gatos lindos en línea!` Has omitido el texto o tienes un error tipográfico.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  ?.innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

Tu elemento `p` debería estar debajo del elemento `h2`. Los tienes en el orden incorrecto.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H2'), collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>
    
--fcc-editable-region--
  </body>
</html>
```
