---
id: 587d824f367417b2b2512c5a
title: Funktionale Tests einer API-Antwort mit Chai-HTTP III ausführen – PUT-Methode
challengeType: 2
forumTopicId: 301590
dashedName: run-functional-tests-on-an-api-response-using-chai-http-iii---put-method
---

# --description--

Zur Erinnerung: Dieses Projekt baut auf dem folgenden Starter-Projekt auf, das von <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

Wenn Sie eine `PUT` Anfrage testen, senden Sie oft Daten mit. Die Daten, die Sie mit Ihrer `PUT` Anfrage mitsenden, nennt man den Body der Anfrage.

Um eine `PUT` Anfrage und ein JSON-Objekt an den `'/travellers'` Endpunkt zu senden, können Sie die `chai-http` Plugin-Methoden `put` und `send` verwenden:

```js
chai
  .request(server)
  .keepOpen()
  .put('/travellers')
  .send({
    "surname": [last name of a traveller of the past]
  })
  ...
```

Und die Route antwortet mit:

```json
{
  "name": [first name],
  "surname": [last name],
  "dates": [birth - death years]
}
```

Sehen Sie sich den Server-Code für die verschiedenen Antworten auf den `'/travellers'` Endpunkt an.

# --instructions--

Ändern Sie innerhalb von `tests/2_functional-tests.js` den `'Send {surname: "Colombo"}'` Test (`// #3`) und verwenden Sie die `put` und `send` Methoden, um den `'/travellers'` Endpunkt zu testen.

Senden Sie das folgende JSON-Objekt mit Ihrer PUT-Anfrage:

```json
{
  "surname": "Colombo"
}
```

Prüfen Sie innerhalb des `request.end` Callbacks Folgendes:

1.  Der `status` sollte `200` sein
2.  Der `type` sollte `application/json` sein
3.  Der `body.name` sollte `Cristoforo` sein
4.  Der `body.surname` sollte `Colombo` sein

Folgen Sie der oben angegebenen Reihenfolge der Zusicherungen – darauf bauen wir. Entfernen Sie außerdem `assert.fail()`, sobald Sie fertig sind.

# --hints--

Alle Tests sollten bestanden werden.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Sie sollten testen, ob `res.status` den Wert 200 hat.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

Sie sollten testen, ob `res.type` `'application/json'` ist.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.type');
assert.match(data.assertions[1].args[1], /('|")application\/json\1/);
```

Sie sollten testen, ob `res.body.name` `'Cristoforo'` ist.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'equal');
assert.equal(data.assertions[2].args[0], 'res.body.name');
assert.match(data.assertions[2].args[1], /('|")Cristoforo\1/);
```

Sie sollten testen, ob `res.body.surname` `'Colombo'` ist.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'equal');
assert.equal(data.assertions[3].args[0], 'res.body.surname');
assert.match(data.assertions[3].args[1], /('|")Colombo\1/);
```

