---
id: 587d824b367417b2b2512c49
title: Test auf Wahrhaftigkeit
challengeType: 2
forumTopicId: 301596
dashedName: test-for-truthiness
---

# --description--

Zur Erinnerung: Dieses Projekt baut auf dem folgenden Starter-Projekt auf, das von <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

`isTrue()` testet auf den Boolean-Wert `true` und `isNotTrue()` besteht, wenn alles außer dem Boolean-Wert `true` übergeben wird.

```js
assert.isTrue(true, 'This will pass with the boolean value true');
assert.isTrue('true', 'This will NOT pass with the string value "true"');
assert.isTrue(1, 'This will NOT pass with the number value 1');
```

`isFalse()` und `isNotFalse()` existieren ebenfalls und verhalten sich ähnlich wie ihre true-Pendants, außer dass sie nach dem Boolean-Wert `false` suchen.

# --instructions--

Ändern Sie innerhalb von `tests/1_unit-tests.js` unter dem Test mit der Bezeichnung `#4` in der `Basic Assertions`-Suite jedes `assert` entweder in `assert.isTrue` oder `assert.isNotTrue`, damit der Test besteht (sollte zu `true` ausgewertet werden). Ändern Sie nicht die Argumente, die an die Zusicherungen übergeben werden.

# --hints--

Alle Tests sollten bestehen.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Sie sollten für die erste Zusicherung die richtige Methode auswählen – `isTrue` vs. `isNotTrue`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isTrue', 'True is true');
```

Sie sollten für die zweite Zusicherung die richtige Methode auswählen – `isTrue` vs. `isNotTrue`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'isTrue',
    'Double negation of a truthy value is true'
  );
```

Sie sollten für die dritte Zusicherung die richtige Methode auswählen – `isTrue` vs. `isNotTrue`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'isNotTrue',
    'A truthy object is not true - neither is a false one'
  );
```

