---
id: 587d824d367417b2b2512c50
title: "Prüfen, ob ein Wert ein Array ist"
challengeType: 2
forumTopicId: 301600
dashedName: test-if-a-value-is-an-array
---

# --description--

Zur Erinnerung: Dieses Projekt baut auf dem folgenden Starter-Projekt auf, das von <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

# --instructions--

Ändern Sie innerhalb von `tests/1_unit-tests.js` im Test mit der Bezeichnung `#11` in der `Arrays`-Suite jeweils `assert` entweder in `assert.isArray` oder `assert.isNotArray`, damit der Test besteht (sollte zu `true` ausgewertet werden). Verändern Sie nicht die Argumente, die an die Zusicherungen übergeben werden.

# --hints--

Alle Tests sollten bestanden werden.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=10');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Sie sollten für die erste Zusicherung die richtige Methode auswählen – `isArray` oder `isNotArray`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=10');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isArray',
  'String.prototype.split() returns an Array'
);
```

Sie sollten für die zweite Zusicherung die richtige Methode auswählen – `isArray` oder `isNotArray`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=10');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'isNotArray',
  'Array.prototype.indexOf() returns a number'
);
```

