---
id: 587d824d367417b2b2512c53
title: Testen, ob ein String einen Substring enthält
challengeType: 2
forumTopicId: 301597
dashedName: test-if-a-string-contains-a-substring
---

# --description--

As a reminder, this project is being built upon the following starter project cloned from <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`include()` und `notInclude()` funktionieren auch für Strings! `include()` überprüft, ob der aktuelle String den erwarteten Substring enthält.

# --instructions--

Ändere innerhalb `tests/1_unit-tests.js` unter dem Test mit der Beschriftung `#14` in der `Strings`-Suite, jeden `assert` in `assert.include` oder `assert.notInclude` um, damit der Test bestanden wird (sollte `true` ausgeben). Ändere nicht die Argumente, die an die Assertions übergeben werden.

# --hints--

Alle Tests sollten bestehen.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Du solltest die richtige Methode für die erste Behauptung auswählen - `include` vs. `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'include',
  "'Arrow' contains 'row'..."
);
```

Du solltest die richtige Methode für die zweite Behauptung auswählen - `include` vs. `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'notInclude',
  "... a 'dart' doesn't contain a 'queue'"
);
```

