---
id: 587d824e367417b2b2512c57
title: "Prüfen, ob ein Objekt eine Instanz eines Konstruktors ist"
challengeType: 2
forumTopicId: 301605
dashedName: test-if-an-object-is-an-instance-of-a-constructor
---

# --description--

Zur Erinnerung: Dieses Projekt baut auf dem folgenden Starter-Projekt auf, das von <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

`#instanceOf` stellt sicher, dass ein Objekt eine Instanz eines Konstruktors ist.

# --instructions--

Ändern Sie innerhalb von `tests/1_unit-tests.js` unter dem Test mit der Bezeichnung `#18` in der `Objects`-Suite jeweils `assert` entweder in `assert.instanceOf` oder `assert.notInstanceOf`, damit der Test besteht (sollte zu `true` auswerten). Verändern Sie nicht die Argumente, die an die Zusicherungen übergeben werden.

# --hints--

Alle Tests sollten bestanden werden.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Sie sollten für die erste Zusicherung die richtige Methode auswählen – `instanceOf` oder `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notInstanceOf',
  'myCar is not an instance of Plane'
);
```

Sie sollten für die zweite Zusicherung die richtige Methode auswählen – `instanceOf` oder `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'instanceOf',
  'airlinePlane is an instance of Plane'
);
```

Sie sollten für die dritte Zusicherung die richtige Methode auswählen – `instanceOf` oder `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[2].method,
  'instanceOf',
  'everything is an Object in JavaScript...'
);
```

Sie sollten für die vierte Zusicherung die richtige Methode auswählen – `instanceOf` oder `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'notInstanceOf',
  'myCar.wheels is not an instance of String'
);
```

