---
id: 587d824d367417b2b2512c51
title: Teste, ob ein Array ein Element enthält
challengeType: 2
forumTopicId: 301603
dashedName: test-if-an-array-contains-an-item
---

# --description--

As a reminder, this project is being built upon the following starter project cloned from <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

Ändere innerhalb `tests/1_unit-tests.js` unter dem Test mit der Beschriftung `#12` in der `Arrays`-Suite, jeden `assert` in `assert.include` oder `assert.notInclude` um, damit der Test bestanden wird (sollte `true` ausgeben). Ändere nicht die Argumente, die an die Assertions übergeben werden.

# --hints--

Alle Tests sollten durchlaufen.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Du solltest die richtige Methode für die erste Behauptung wählen - `include` vs. `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notInclude',
  "It's summer in july..."
);
```

Du solltest die richtige Methode für die zweite Behauptung auswählen - `include` vs. `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'include',
  'JavaScript is a backend language !!'
);
```

