---
id: 587d824f367417b2b2512c5b
title: Funktionale Tests einer API-Antwort mit Chai-HTTP IV ausführen – PUT-Methode
challengeType: 2
forumTopicId: 301591
dashedName: run-functional-tests-on-an-api-response-using-chai-http-iv---put-method
---

# --description--

Zur Erinnerung: Dieses Projekt baut auf dem folgenden Starter-Projekt auf, das von <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

Diese Übung ähnelt der vorherigen.

Da Sie nun wissen, wie man eine `PUT` Anfrage testet, sind Sie jetzt dran, dies von Grund auf selbst zu machen.

# --instructions--

Ändern Sie innerhalb von `tests/2_functional-tests.js` den `'Send {surname: "da Verrazzano"}'` Test (`// #4`) und verwenden Sie die `put` und `send` Methoden, um den `'/travellers'` Endpunkt zu testen.

Senden Sie mit Ihrer PUT-Anfrage das folgende JSON-Objekt:

```json
{
  "surname": "da Verrazzano"
}
```

Prüfen Sie im `request.end` Callback Folgendes:

1.  Der `status` sollte `200` sein
2.  Der `type` sollte `application/json` sein
3.  Der `body.name` sollte `Giovanni` sein
4.  Der `body.surname` sollte `da Verrazzano` sein

Folgen Sie der oben angegebenen Reihenfolge der Zusicherungen – darauf bauen wir. Entfernen Sie außerdem `assert.fail()`, sobald Sie fertig sind.

# --hints--

Alle Tests sollten bestanden werden

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Sie sollten testen, ob `res.status` 200 ist

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

Sie sollten testen, ob `res.type` `'application/json'` ist

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.type');
assert.match(data.assertions[1].args[1], /('|")application\/json\1/);
```

Sie sollten testen, ob `res.body.name` `'Giovanni'` ist

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'equal');
assert.equal(data.assertions[2].args[0], 'res.body.name');
assert.match(data.assertions[2].args[1], /('|")Giovanni\1/);
```

Sie sollten testen, ob `res.body.surname` `'da Verrazzano'` ist

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'equal');
assert.equal(data.assertions[3].args[0], 'res.body.surname');
assert.match(data.assertions[3].args[1], /('|")da Verrazzano\1/);
```

