---
id: 587d824d367417b2b2512c54
title: "Reguläre Ausdrücke verwenden, um einen String zu testen"
challengeType: 2
forumTopicId: 301608
dashedName: use-regular-expressions-to-test-a-string
---

# --description--

Zur Erinnerung: Dieses Projekt baut auf dem folgenden Starter-Projekt auf, das von <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

`match()` sichert zu, dass der tatsächliche Wert mit dem regulären Ausdruck des zweiten Arguments übereinstimmt.

# --instructions--

Ändern Sie innerhalb von `tests/1_unit-tests.js` unter dem Test mit der Bezeichnung `#15` in der `Strings`-Suite jeweils `assert` entweder in `assert.match` oder `assert.notMatch`, damit der Test besteht (sollte zu `true` auswerten). Verändern Sie nicht die Argumente, die an die Zusicherungen übergeben werden.

# --hints--

Alle Tests sollten bestehen.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=14');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Sie sollten für die erste Zusicherung die richtige Methode auswählen – `match` oder `notMatch`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=14');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'match',
  "'# name:John Doe, age:35' matches the regex"
);
```

Sie sollten für die zweite Zusicherung die richtige Methode auswählen – `match` oder `notMatch`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=14');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'notMatch',
  "'# name:Paul Smith III, age:twenty-four' does not match the regex (the age must be numeric)"
);
```

