---
id: 587d824e367417b2b2512c56
title: "Testen, ob ein Wert einem bestimmten Datentyp einer Datenstruktur entspricht"
challengeType: 2
forumTopicId: 301601
dashedName: test-if-a-value-is-of-a-specific-data-structure-type
---

# --description--

Zur Erinnerung: Dieses Projekt baut auf dem folgenden Starter-Projekt auf, das von <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

`#typeOf` sichert zu, dass der Typ von value dem angegebenen String entspricht, wie von `Object.prototype.toString` bestimmt.

# --instructions--

Ändern Sie innerhalb von `tests/1_unit-tests.js` unter dem Test mit der Bezeichnung `#17` in der `Objects`-Suite jeweils `assert` in `assert.typeOf` oder `assert.notTypeOf`, damit der Test besteht (sollte zu `true` auswerten). Verändern Sie nicht die Argumente, die an die Zusicherungen übergeben werden.

# --hints--

Alle Tests sollten bestanden werden.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Sie sollten für die erste Zusicherung die richtige Methode auswählen – `typeOf` oder `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'typeOf',
  'myCar is typeOf Object'
);
```

Sie sollten für die zweite Zusicherung die richtige Methode auswählen – `typeOf` oder `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'typeOf',
  'Car.model is a String'
);
```

Sie sollten für die dritte Zusicherung die richtige Methode auswählen – `typeOf` oder `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[2].method,
  'notTypeOf',
  'Plane.wings is a Number (not a String)'
);
```

Sie sollten für die vierte Zusicherung die richtige Methode auswählen – `typeOf` oder `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'typeOf',
  'Plane.engines is an Array'
);
```

Sie sollten für die fünfte Zusicherung die richtige Methode auswählen – `typeOf` oder `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[4].method,
  'typeOf',
  'Car.wheels is a Number'
);
```

