---
id: 587d824c367417b2b2512c4d
title: Eigenschaften von zwei Elementen vergleichen
challengeType: 2
forumTopicId: 301588
dashedName: compare-the-properties-of-two-elements
---

# --description--

Zur Erinnerung: Dieses Projekt baut auf dem folgenden Starter-Projekt auf, das von <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

# --instructions--

Innerhalb von `tests/1_unit-tests.js` unter dem Test mit der Bezeichnung `#8` in der `Comparisons`-Suite ändern Sie jeweils `assert` in `assert.isAbove` oder `assert.isAtMost`, damit der Test besteht (sollte zu `true` auswerten). Ändern Sie nicht die Argumente, die an die Zusicherungen übergeben werden.

# --hints--

Alle Tests sollten bestehen.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Sie sollten für die erste Zusicherung die richtige Methode auswählen – `isAbove` oder `isAtMost`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isAtMost',
  '5 is at most (<=) 5'
);
```

Sie sollten für die zweite Zusicherung die richtige Methode auswählen – `isAbove` oder `isAtMost`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'isAbove', '1 is greater than 0');
```

Sie sollten für die dritte Zusicherung die richtige Methode auswählen – `isAbove` oder `isAtMost`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[2].method,
  'isAbove',
  'Math.PI = 3.14159265 is greater than 3'
);
```

Sie sollten für die vierte Zusicherung die richtige Methode auswählen – `isAbove` oder `isAtMost`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'isAtMost',
  '1 - Math.random() is > 0 and <= 1. It is atMost 1 !'
);
```

