---
id: 587d824f367417b2b2512c59
title: Funktionale Tests für API-Endpunkte mit Chai-HTTP II ausführen
challengeType: 2
forumTopicId: 301592
dashedName: run-functional-tests-on-api-endpoints-using-chai-http-ii
---

# --description--

Zur Erinnerung: Dieses Projekt baut auf dem folgenden Starter-Projekt auf, das von <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

# --instructions--

Ändern Sie innerhalb von `tests/2_functional-tests.js` den `'Test GET /hello with your name'`-Test (`// #2`), um die `status` und die `text` der Antwort zu überprüfen, damit der Test besteht.

Senden Sie Ihren Namen als URL-Abfrage, indem Sie `?name=<your_name>` an die Route anhängen. Der Endpunkt antwortet mit `'hello <your_name>'`.

# --hints--

Alle Tests sollten bestanden werden

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=1');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Sie sollten testen, ob `res.status` == 200 ist

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=1');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

Sie sollten testen, ob `res.text` == `'hello <your_name>'` ist

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=1');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.text');
assert.match(data.assertions[1].args[1], /hello [\w\d_-]/);
```

