---
id: 587d824c367417b2b2512c4f
title: "Testen, ob ein Wert in einem bestimmten Bereich liegt"
challengeType: 2
forumTopicId: 301598
dashedName: test-if-a-value-falls-within-a-specific-range
---

# --description--

Zur Erinnerung: Dieses Projekt baut auf dem folgenden Starter-Projekt auf, das von <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

```javascript
.approximately(actual, expected, delta, [message])
```

Stellt sicher, dass der `actual` innerhalb eines Bereichs von +/- `expected` gleich `delta` ist.

# --instructions--

Ändern Sie innerhalb von `tests/1_unit-tests.js` unter dem Test mit der Bezeichnung `#10` in der `Comparisons`-Suite jeweils jeden `assert` in `assert.approximately`, damit der Test besteht (sollte zu `true` auswerten).

Wählen Sie das Minimum des Bereichs (3. Parameter), damit der Test immer besteht. Es sollte kleiner als 1 sein.

# --hints--

Alle Tests sollten bestehen.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Sie sollten den korrekten Bereich für die erste Zusicherung wählen – `approximately(actual, expected, range)`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'approximately');
assert.equal(
  data.assertions[0].args[2],
  0.5,
  "weirdNumbers(0.5) is in the range (0.5, 1.5]. It's within 1 +/- 0.5"
);
```

Sie sollten den korrekten Bereich für die zweite Zusicherung wählen – `approximately(actual, expected, range)`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'approximately');
assert.equal(
  data.assertions[1].args[2],
  0.8,
  "weirdNumbers(0.2) is in the range (0.2, 1.2]. It's within 1 +/- 0.8"
);
```

