---
id: 587d824b367417b2b2512c4b
title: "Verwenden Sie das dreifache Gleichheitszeichen, um strikte Gleichheit zuzusichern"
challengeType: 2
forumTopicId: 301610
dashedName: use-the-triple-equals-to-assert-strict-equality
---

# --description--

Zur Erinnerung: Dieses Projekt baut auf dem folgenden Starter-Projekt auf, das von <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

`strictEqual()` vergleicht Objekte mit `===`.

# --instructions--

Ändern Sie innerhalb von `tests/1_unit-tests.js` unter dem Test mit der Bezeichnung `#6` in der `Equality`-Suite jedes `assert` entweder in `assert.strictEqual` oder `assert.notStrictEqual`, damit der Test besteht (sollte zu `true` auswerten). Verändern Sie nicht die Argumente, die an die Zusicherungen übergeben werden.

# --hints--

Alle Tests sollten bestanden werden.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Sie sollten für die erste Zusicherung die richtige Methode auswählen – `strictEqual` vs. `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'notStrictEqual',
    'with strictEqual the type must match'
  );
```

Sie sollten für die zweite Zusicherung die richtige Methode auswählen – `strictEqual` vs. `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'strictEqual', '3*2 = 6...');
```

Sie sollten für die dritte Zusicherung die richtige Methode auswählen – `strictEqual` vs. `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'strictEqual',
    "6 * '2' is 12. Types match !"
  );
```

Sie sollten für die vierte Zusicherung die richtige Methode auswählen – `strictEqual` vs. `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[3].method,
    'notStrictEqual',
    'Even if they have the same elements, the Arrays are notStrictEqual'
  );
```

