---
id: 587d824e367417b2b2512c58
title: Funktionale Tests für API-Endpunkte mit Chai-HTTP ausführen
challengeType: 2
forumTopicId: 301593
dashedName: run-functional-tests-on-api-endpoints-using-chai-http
---

# --description--

Zur Erinnerung: Dieses Projekt basiert auf dem folgenden Starter-Projekt, das von <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

Mocha ermöglicht es Ihnen, asynchrone Operationen wie Aufrufe von API-Endpunkten mit einem Plugin namens `chai-http` zu testen.

Das Folgende ist ein Beispiel für einen Test mit `chai-http` für eine Suite namens `'GET /hello?name=[name] => "hello [name]"'`:

```js
suite('GET /hello?name=[name] => "hello [name]"', function () {
  test('?name=John', function (done) {
    chai
      .request(server)
      .keepOpen()
      .get('/hello?name=John')
      .end(function (err, res) {
        assert.equal(res.status, 200, 'Response status should be 200');
        assert.equal(res.text, 'hello John', 'Response should be "hello John"');
        done();
      });
  });
});
```

Der Test sendet eine `GET`-Anfrage an den Server mit einem Namen als URL-Abfragezeichenfolge (`?name=John`). In der Callback-Funktion der `end`-Methode wird das Antwortobjekt (`res`) empfangen, das die `status`-Eigenschaft enthält.

Die erste `assert.equal` prüft, ob der Status gleich `200` ist. Die zweite `assert.equal` überprüft, dass der Antwortstring (`res.text`) gleich `"hello John"` ist.

Beachten Sie auch den `done`-Parameter in der Callback-Funktion des Tests. Es ist notwendig, ihn am Ende eines Tests ohne Argument aufzurufen, um zu signalisieren, dass die asynchrone Operation abgeschlossen ist.

Schließlich beachten Sie die `keepOpen`-Methode direkt nach der `request`-Methode. Normalerweise würden Sie Ihre Tests über die Befehlszeile ausführen oder als Teil eines automatisierten Integrationsprozesses, und Sie könnten `chai-http` verwenden, um Ihren Server automatisch zu starten und zu stoppen.

Die Tests, die beim Absenden des Links zu Ihrem Projekt ausgeführt werden, erfordern jedoch, dass Ihr Server läuft, daher müssen Sie die `keepOpen`-Methode verwenden, um zu verhindern, dass `chai-http` Ihren Server stoppt.

# --instructions--

Ändern Sie innerhalb von `tests/2_functional-tests.js` den `'Test GET /hello with no name'`-Test (`// #1`), um die `status` und die `text` der Antwort zu überprüfen, damit der Test besteht. Ändern Sie nicht die Argumente, die an die Assertions übergeben werden.

Es sollte keine URL-Abfrage vorhanden sein. Ohne eine Namens-URL-Abfrage antwortet der Endpunkt mit `hello Guest`.

# --hints--

Alle Tests sollten bestehen

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Sie sollten testen, ob `res.status` == 200 ist

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

Sie sollten testen, ob `res.text` == `'hello Guest'` ist

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.text');
assert.match(data.assertions[1].args[1], /('|")hello Guest\1/);
```

