---
id: 587d824d367417b2b2512c52
title: "Prüfen, ob ein Wert ein String ist"
challengeType: 2
forumTopicId: 301599
dashedName: test-if-a-value-is-a-string
---

# --description--

Zur Erinnerung: Dieses Projekt baut auf dem folgenden Starter-Projekt auf, das von <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

`isString` oder `isNotString` sichert zu, dass der tatsächliche Wert ein String ist.

# --instructions--

Ändern Sie innerhalb von `tests/1_unit-tests.js` unter dem Test mit der Bezeichnung `#13` in der `Strings`-Suite jeweils `assert` entweder in `assert.isString` oder `assert.isNotString`, damit der Test besteht (sollte zu `true` auswerten). Verändern Sie nicht die Argumente, die an die Zusicherungen übergeben werden.

# --hints--

Alle Tests sollten bestanden werden.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Sie sollten für die erste Zusicherung die richtige Methode auswählen – `isString` oder `isNotString`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isNotString',
  'A float number is not a string'
);
```

Sie sollten für die zweite Zusicherung die richtige Methode auswählen – `isString` oder `isNotString`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'isString',
  'environment vars are strings (or undefined)'
);
```

Sie sollten für die dritte Zusicherung die richtige Methode auswählen – `isString` oder `isNotString`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'isString', 'A JSON is a string');
```

