---
id: 5f8884f4c46685731aabfc41
title: Funktionale Tests mit einem Headless-Browser ausführen II
challengeType: 2
forumTopicId: 301594
dashedName: run-functional-tests-using-a-headless-browser-ii
---

# --description--

Zur Erinnerung: Dieses Projekt baut auf dem folgenden Starter-Projekt auf, das von <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

# --instructions--

Innerhalb von `tests/2_functional-tests.js`, im `'Submit the surname "Vespucci" in the HTML form'`-Test (`// #6`), automatisieren Sie Folgendes:

1.  Füllen Sie das Formular mit dem Nachnamen `Vespucci` aus
2.  Drücken Sie die Absenden-Schaltfläche

Und innerhalb des `pressButton`-Callbacks:

1.  Überprüfen Sie, dass der Status OK ist `200`
2.  Überprüfen Sie, dass der Text im Element `span#name` `'Amerigo'` ist
3.  Überprüfen Sie, dass der Text im Element `span#surname` `'Vespucci'` ist
4.  Überprüfen Sie, dass das/die Element(e) `span#dates` existieren und deren Anzahl `1` beträgt

Vergessen Sie nicht, den `assert.fail()`-Aufruf zu entfernen.

# --hints--

Alle Tests sollten bestanden werden.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Sie sollten überprüfen, dass die Anfrage des Headless-Browsers erfolgreich war.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'browser.success');
```

Sie sollten überprüfen, dass der Text im Element `span#name` `'Amerigo'` ist.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'browser.text');
assert.match(data.assertions[1].args[0], /('|")span#name\1/);
assert.match(data.assertions[1].args[1], /('|")Amerigo\1/);
```

Sie sollten überprüfen, dass der Text im Element `span#surname` `'Vespucci'` ist.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'browser.text');
assert.match(data.assertions[2].args[0], /('|")span#surname\1/);
assert.match(data.assertions[2].args[1], /('|")Vespucci\1/);
```

Sie sollten überprüfen, dass das Element `span#dates` existiert und dessen Anzahl 1 beträgt.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'browser.elements');
assert.match(data.assertions[3].args[0], /('|")span#dates\1/);
assert.equal(data.assertions[3].args[1], 1);
```

