---
id: 69de4308e652aaeb5ccdb03e
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika warsha hii, utafanya mazoezi ya kutumia njia za mfuatano wa herufi za Python.

Anza kwa kuunda kigezo kinachoitwa `email` na ukae thamani ya mfuatano wa herufi `alice.johnson@company.com`.

# --hints--

Unapaswa kuunda kigezo kinachoitwa `email`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email')
`) })
```

Unapaswa kuweka thamani ya mfuatano wa herufi `alice.johnson@company.com` katika `email`. Kumbuka kufunga maandishi hayo kwa kutumia alama za nukuu moja au alama za nukuu maradufu.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('email').is_equivalent("email = 'alice.johnson@company.com'")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
