---
id: 69e20876baed33de636b989a
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Wakati mwingine, unaweza kutaka kupata nafasi ya herufi maalum ndani ya mfuatano wa herufi. Python hutoa njia iliyo jengwa ndani `.find()` kwa ajili ya kusudi hili. Inarudisha nambari kamili inayowakilisha kielezo (nafasi) ya tukio la kwanza la mfuatano wa herufi uliotolewa.

Hapa kuna mfano:

```py
text = 'hello'
print(text.find('e'))  # 1
```

Sasa, fanya wito wa njia `.find()` kwenye `email` na uipe `'@'` ili kupata nafasi ya alama ya at katika mfuatano wa herufi `email`. Weka thamani hiyo kwenye kigezo kinachoitwa `at_position`.

Kisha chapisha `at_position` ili matokeo yaoneshe: `Position of @: 13`.

# --hints--

Unapaswa kuunda kigezo kinachoitwa `at_position`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_position')
`) })
```

Unapaswa kuweka thamani ya `email.find('@')` kwenye `at_position`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_position').is_equivalent("at_position = email.find('@')")
`) })
```

Unapaswa kuchapisha `at_position` ili matokeo yaoneshe `Position of @: 13`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Position of @: {at_position}')",
    "print('Position of @: ' + str(at_position))",
    "print('Position of @:', at_position)",
    "print('Position of @:', str(at_position))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')
--fcc-editable-region--

--fcc-editable-region--
```
