---
id: 69e20a0451c7a3ef189514d3
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Sasa unajua nafasi ya `@` katika barua pepe, unaweza kutoa jina la mtumizi, ambalo lipo kabla ya alama ya `@`.

Katika warsha iliyopita, ulitoa sehemu ya mfuatano wa herufi kwa kutumia slicing. Kama ukumbusho, sintaksia ni:

```python
text = 'hello world'
print(text[:5])  # hello
```

Tumia slicing kwenye `email` kutoa kila kitu kabla ya `at_position`, na hifadhi matokeo katika kigezo kinachoitwa `username`.

Kisha chapisha `username` ili matokeo yaoneshe `Username: alice.johnson`.

# --hints--

Unapaswa kuunda kigezo kinachoitwa `username`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('username')
`) })
```

Unapaswa kutumia slicing na `at_position` kwenye `email` kuweka thamani kwa `username`.

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('username')
assert node.is_equivalent("username = email[:at_position]") or node.is_equivalent("username = email[0:at_position]")
`) })
```

Unapaswa kuchapisha `username` ili matokeo yaoneshe `Username: alice.johnson`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {username}')",
    "print('Username: ' + username)",
    "print('Username:', username)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

--fcc-editable-region--

--fcc-editable-region--
```
