---
id: 69e62ce96df237884da25282
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Njia `.lower()` hubadilisha herufi zote katika mfuatano wa herufi kuwa herufi ndogo. Hii ni muhimu kwa kusawazisha anwani za barua pepe:

```python
upperGreet = 'HELLO WORLD'
lowerGreet = upperGreet.lower()
print(lowerGreet)  # Output: hello world
```

Unaweza kuita `.strip()` na `.lower()` kwa mfuatano wa herufi kwa njia hii:

```python
messy = '  HELLO WORLD  '
clean = messy.strip().lower()
print(clean)  # Output: hello world
```

Tengeneza kigezo `messy_email` na weka thamani ya mfuatano wa herufi `SARAH.DAVIS@COMPANY.COM` ukiwa na nafasi mbili mwanzoni na nafasi mbili mwishoni.

Pia, wito `.strip().lower()` kwenye `messy_email` na weka matokeo katika kigezo `clean_email`.

Kisha chapisha zote mbili ili matokeo yawe kama ifuatavyo:

```md
Original email: '  SARAH.DAVIS@COMPANY.COM  '
Cleaned email: 'sarah.davis@company.com'
```

# --hints--

Unapaswa kutengeneza kigezo kinachoitwa `messy_email`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_email')
`) })
```

Unapaswa kuweka <code>'&nbsp;&nbsp;SARAH.DAVIS@COMPANY.COM&nbsp;&nbsp;'</code> katika `messy_email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_email').is_equivalent("messy_email = '  SARAH.DAVIS@COMPANY.COM  '")
`) })
```

Unapaswa kutengeneza kigezo kinachoitwa `clean_email`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('clean_email')
`) })
```

Unapaswa kuita `.strip().lower()` kwenye `messy_email` na weka matokeo katika kigezo `clean_email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_email').is_equivalent("clean_email = messy_email.strip().lower()")
`) })
```

Unapaswa kuchapisha `messy_email` ili matokeo yawe `Original email: '  SARAH.DAVIS@COMPANY.COM  '`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original email: '{messy_email}'\\")",
    "print(\\"Original email: '\\" + messy_email + \\"'\\")",
    "print(\\"Original email: '\\", messy_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

Unapaswa kuchapisha `clean_email` ili matokeo yawe `Cleaned email: 'sarah.davis@company.com'`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned email: '{clean_email}'\\")",
    "print(\\"Cleaned email: '\\" + clean_email + \\"'\\")",
    "print(\\"Cleaned email: '\\", clean_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
