---
id: 69e6299710319e7cb8ebdc37
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Data halisi mara nyingi huwa na machache ya ziada ya nafasi na matumizi yasiyo thabiti ya herufi kubwa na ndogo.

Tengeneza mabadiliko `messy_name` na weka thamani ya mfuatano wa herufi `sARaH dAVis` yenye nafasi mbili mwanzoni na nafasi mbili mwishoni. Kisha chapisha ili matokeo yawe:

```md
Original name: '  sARaH dAVis  '
```

Utayosafisha haya katika hatua zinazofuata.

# --hints--

Unapaswa kutengeneza mabadiliko yenye jina `messy_name`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_name')
`) })
```

Unapaswa kuweka thamani <code>'&nbsp;&nbsp;sARaH dAVis&nbsp;&nbsp;'</code> kwa `messy_name`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_name').is_equivalent("messy_name = '  sARaH dAVis  '")
`) })
```

Unapaswa kuchapisha `messy_name` ili matokeo yawe `Original name: '  sARaH dAVis  '`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original name: '{messy_name}'\\")",
    "print(\\"Original name: '\\" + messy_name + \\"'\\")",
    "print(\\"Original name: '\\", messy_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

--fcc-editable-region--

--fcc-editable-region--
```
