---
id: 69e6284f3b7659c5a05c2a79
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Unaweza kuangalia kama sehemu ndogo ya mfuatano wa herufi ipo ndani ya mfuatano wa herufi kwa kutumia neno kuu `in`. Inarudisha `True` ikiwa sehemu ndogo ipo na `False` vinginevyo. Hapa kuna mfano:

```py
sentence = 'I love Python'
print('Python' in sentence) # True
print('Java' in sentence) # False
```

Tengeneza kigezo kinachoitwa `has_at` kinachokagua kama herufi `@` ipo katika `email`.

Kisha chapisha `has_at` ili matokeo yaoneshe: `Has @ symbol: True`.

# --hints--

Unapaswa kutengeneza kigezo kinachoitwa `has_at`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_at')
`) })
```

Kigezo chako cha `has_at` kinapaswa kuwa na thamani `'@' in email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_at').is_equivalent("has_at = '@' in email")
`) })
```

Unapaswa kuchapisha `has_at` ili matokeo yaoneshe `Has @ symbol: True`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has @ symbol: {has_at}')",
    "print('Has @ symbol: ' + str(has_at))",
    "print('Has @ symbol:', has_at)",
    "print('Has @ symbol:', str(has_at))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

--fcc-editable-region--

--fcc-editable-region--
```
