---
id: 69e76e91402d7850510a5f18
title: Hatua 16
challengeType: 20
dashedName: step-16
---

# --description--

Njia `.join(list)` ni kinyume cha `.split()`. Inachanganya orodha ya mfuatano wa herufi kuwa mfuatano mmoja kwa kutumia kitenzi cha kutenganisha. Kwa mfano:

```py
parts = ['apple', 'banana', 'cherry']
joined = ' - '.join(parts)
print(joined)  # Output: apple - banana - cherry
```

Unganisha vipengele kutoka `address_parts` kwa kutumia tabia ya `|` iliyozungukwa na nafasi na weka thamani hiyo katika kigezo kinachoitwa `rejoined`.

Kisha chapisha `rejoined` ili matokeo ya chapisho yawe: `Rejoined: 123 Main Street | Springfield | IL`.

# --hints--

Unapaswa kuunda kigezo kinachoitwa `rejoined`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('rejoined')
`) })
```

Unapaswa kuweka thamani ya `rejoined` kuwa `' | '.join(address_parts)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('rejoined').is_equivalent("rejoined = ' | '.join(address_parts)")
`) })
```

Unapaswa kuchapisha `rejoined` ili matokeo ya chapisho yawe `Rejoined: 123 Main Street | Springfield | IL`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Rejoined: {rejoined}')",
    "print('Rejoined: ' + rejoined)",
    "print('Rejoined:', rejoined)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print(f'Address parts: {address_parts}')

--fcc-editable-region--

--fcc-editable-region--
```
