---
id: 69e736c39dfa403496e8071e
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Njia `.upper()` hubadilisha herufi zote kuwa herufi kubwa. Hii ni muhimu kwa alama za majina au vichwa vya sehemu:

```py
name = 'sarah davis'
badge = name.upper()
print(badge)  # Output: SARAH DAVIS
```

Tengeneza kigezo kinachoitwa `name_badge` na weka thamani yake kuwa matokeo ya wito wa njia `.upper()` kwenye `clean_name`.

Kisha chapisha `name_badge` ili matokeo yawe: `Name badge: SARAH DAVIS`.

# --hints--

Unapaswa kutengeneza kigezo kinachoitwa `name_badge`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('name_badge')
`) })
```

Unapaswa kuita `.upper()` kwenye `clean_name` na weka thamani ya matokeo kwenye kigezo `name_badge`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('name_badge').is_equivalent("name_badge = clean_name.upper()")
`) })
```

Unapaswa kuchapisha `name_badge` ili matokeo yawe `Name badge: SARAH DAVIS`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Name badge: {name_badge}')",
    "print('Name badge: ' + name_badge)",
    "print('Name badge:', name_badge)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

--fcc-editable-region--

--fcc-editable-region--
```
