---
id: 69e73f4f5f986657647bbd63
title: Hatua 13
challengeType: 20
dashedName: step-13
---

# --description--

Njia ya `.replace(old, new)` huweka thamani kila tukio la sehemu ndogo ya mfuatano wa herufi `old` na nyingine, `new`. Pia unaweza kuunganisha na njia nyingine. Kwa mfano:

```py
text = 'john.doe'
name = text.replace('.', ' ').title()
print(name)  # Output: John Doe
```

Badilisha `.` kwa nafasi moja katika `email_user`, kisha fanya wito wa `.title()` kwenye matokeo. Weka thamani hiyo katika kigezo kinachoitwa `display_name`.

Kisha chapisha `display_name` ili matokeo yaoneshe: `Display name: Sarah Davis`.

# --hints--

Unapaswa kuunda kigezo kinachoitwa `display_name`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('display_name')
`) })
```

Unapaswa kutumia `.replace('.', ' ')` na `.title()` kwenye `email_user` na kuweka thamani hiyo katika kigezo cha `display_name`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('display_name').is_equivalent("display_name = email_user.replace('.', ' ').title()")
`) })
```

Unapaswa kuchapisha `display_name` ili matokeo yaoneshe `Display name: Sarah Davis`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Display name: {display_name}')",
    "print('Display name: ' + display_name)",
    "print('Display name:', display_name)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

--fcc-editable-region--

--fcc-editable-region--
```
