---
id: 69e76d1997661be7f0040e2d
title: Hatua 15
challengeType: 20
dashedName: step-15
---

# --description--

Njia `.split(separator)` hugawanya mfuatano wa herufi kuwa orodha ya sehemu ndogo ndogo kila inapokutana na `separator`. Kwa mfano:

```py
fruits = 'apple, banana, cherry'
parts = fruits.split(', ')
print(parts)  # Output: ['apple', 'banana', 'cherry']
```

Tengeneza kigezo `full_address` na weka thamani yake kuwa mfuatano wa herufi `123 Main Street, Springfield, IL`. Kisha tengeneza `address_parts` kwa kugawanya kwa koma ikifuatiwa na nafasi moja.

Kisha chapisha `address_parts` ili matokeo yaoneshe `Address parts: ['123 Main Street', 'Springfield', 'IL']`.

# --hints--

Unapaswa kutengeneza kigezo kinachoitwa `full_address`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('full_address')
`) })
```

Unapaswa kuweka thamani ya mfuatano wa herufi `123 Main Street, Springfield, IL` kwa `full_address`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('full_address').is_equivalent("full_address = '123 Main Street, Springfield, IL'")
`) })
```

Unapaswa kutengeneza kigezo kinachoitwa `address_parts`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('address_parts')
`) })
```

Unapaswa kutumia `.split(', ')` kwenye `full_address` na weka matokeo kwa kigezo `address_parts`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('address_parts').is_equivalent("address_parts = full_address.split(', ')")
`) })
```

Unapaswa kuchapisha `address_parts` ili matokeo yaoneshe `Address parts: ['123 Main Street', 'Springfield', 'IL']`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Address parts: {address_parts}')",
    "print('Address parts: ' + str(address_parts))",
    "print('Address parts:', address_parts)",
    "print('Address parts:', str(address_parts))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

--fcc-editable-region--

--fcc-editable-region--
```
