---
id: 69e73a141addc6f442924b0d
title: Hatua 12
challengeType: 20
dashedName: step-12
---

# --description--

Umeweza kutumia `.find()` na kukata sehemu kwenye `email` tayari. Sasa utatumia mbinu ile ile kwenye `clean_email`.

Tafuta nafasi ya herufi `@` katika `clean_email` na uhifadhi katika `at_pos`. Kisha tumia kukata sehemu kuchukua kila kitu kabla ya `@` na uhifadhi katika `email_user`.

Kisha chapisha `clean_email` na `email_user` ili matokeo ya chapisho yawe kama ifuatavyo:

```md
Email: sarah.davis@company.com
Username: sarah.davis
```

# --hints--

Unapaswa kuunda kigezo kinachoitwa `at_pos`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_pos')
`) })
```

Unapaswa kuita `.find('@')` kwenye `clean_email` na weka matokeo katika kigezo `at_pos`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_pos').is_equivalent("at_pos = clean_email.find('@')")
`) })
```

Unapaswa kuunda kigezo kinachoitwa `email_user`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email_user')
`) })
```

Unapaswa kutumia kukata sehemu kwa `at_pos` kwenye `clean_email` na weka matokeo katika kigezo `email_user`.

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('email_user')
assert node.is_equivalent("email_user = clean_email[:at_pos]") or node.is_equivalent("email_user = clean_email[0:at_pos]")
`) })
```

Unapaswa kuchapisha `clean_email` ili matokeo ya chapisho yawe `Email: sarah.davis@company.com`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {clean_email}')",
    "print('Email: ' + clean_email)",
    "print('Email:', clean_email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

Unapaswa kuchapisha `email_user` ili matokeo ya chapisho yawe `Username: sarah.davis`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {email_user}')",
    "print('Username: ' + email_user)",
    "print('Username:', email_user)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

--fcc-editable-region--

--fcc-editable-region--
```
