---
id: 6a71b3e49668c3be461257b5
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Njia ya `.title()` huweka herufi ya kwanza ya kila neno kuwa kubwa. Hapa kuna mfano:

```py
lowercase = 'hello world'
title_case = lowercase.title()
print(title_case)  # Output: Hello World
```

Tengeneza kigezo kinachoitwa `clean_name` na weka thamani yake kuwa matokeo ya wito wa `.title()` kwenye `messy_name`.

# --hints--

Unapaswa kutengeneza kigezo kinachoitwa `clean_name`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('clean_name')
`) })
```

Unapaswa kutumia `.title()` kwenye `messy_name` ili kuweka thamani kwa `clean_name`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_name').is_equivalent("clean_name = messy_name.title()")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
