---
id: 6a86b9e64da746838e67fb13
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

Njia ya `.strip()` huondoa nafasi tupu kutoka pande zote mbili za mfuatano wa herufi.

```python
messy = '  hello world  '
clean = messy.strip()
print(clean)  # Output: hello world
```

Unaweza kuunganisha `.strip()` na `.title()` kuondoa nafasi tupu kutoka pande zote mbili za mfuatano wa herufi na kuweka herufi ya kwanza ya kila neno kuwa kubwa:

```python
messy = '  hello world  '
clean = messy.strip().title()
print(clean)  # Output: Hello World
```

Sahihisha thamani ya `clean_name` kwa kufanya wito wa `.strip()` kwenye `messy_name` kabla ya `.title()`.

Baada ya hapo, chapisha `clean_name` ili matokeo yawe `Cleaned name: 'Sarah Davis'`.

# --hints--

Unapaswa kusahihisha `clean_name` kwa kufanya wito wa `.strip()` kwenye `messy_name` kabla ya `.title()`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_name').is_equivalent("clean_name = messy_name.strip().title()")
`) })
```

Unapaswa kuchapisha `clean_name` ili matokeo yawe `Cleaned name: 'Sarah Davis'`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned name: '{clean_name}'\\")",
    "print(\\"Cleaned name: '\\" + clean_name + \\"'\\")",
    "print(\\"Cleaned name: '\\", clean_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--
clean_name = messy_name.title()
--fcc-editable-region--
```
