---
id: 6a870cc60c23f0332a2fe788
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Kumbuka kutoka kwenye warsha iliyopita kwamba unaweza kutumia f-string ili kujumuisha kigezo ndani ya mfuatano wa herufi, kwa kuandika `f` kabla ya alama ya kunukuu ya mwanzo na kufunga kigezo ndani ya `{}`.

Hapa kuna ukumbusho wa jinsi inavyofanya kazi:

```python
name = 'Alice'
print(f'Hello, {name}')  # Hello, Alice
```

Sasa, chapisha `email` ili matokeo yaoneshe: `Email: alice.johnson@company.com`.

# --hints--

Unapaswa kuchapisha `email` ili matokeo yaoneshe `Email: alice.johnson@company.com`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {email}')",
    "print('Email: ' + email)",
    "print('Email:', email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
--fcc-editable-region--

--fcc-editable-region--
```
