---
id: 69b279bae0f20eda63163cf9
title: 15단계
challengeType: 1
dashedName: step-15
---

# --description--

생성한 `while` 반복문은 `leftPart`의 한 요소와 `rightPart`의 다른 요소를 비교한 다음, 더 작은 요소를 메인 `array` 리스트에 추가합니다.

비교할 요소가 더 이상 없을 때까지 이 작업을 계속합니다. 하지만 `leftPart`에는 여전히 요소가 남아 있고 `rightPart`에는 없을 수도 있고, 그 반대일 수도 있습니다.

`while`이 `leftArrayIndex`의 길이보다 작을 때 실행되는 또 다른 `leftPart` 반복문을 만드세요. 다음 단계에서 `leftPart`에 남아 있는 요소들을 `array` 리스트에 복사하는 데 사용할 것입니다.

# --hints--

조건이 `while`인 `leftArrayIndex < leftPart.length` 반복문을 만들어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
