---
id: 69b279bae0f20eda63163d03
title: 25단계
challengeType: 1
dashedName: step-25
---

# --description--

`numbers`가 정렬되었으니, `'Sorted array: '`을 콘솔에 출력한 다음 `numbers` 배열을 출력하세요.

이로써 병합 정렬 알고리즘이 완성됩니다.

# --hints--

`'Sorted array: '`을 콘솔에 출력해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\.log\s*\(\s*['"]Sorted array:\s*['"]\s*\)\s*;?/
);
```

`numbers`을 출력한 후 `'Sorted array: '` 배열을 콘솔에 출력해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(code),
  /mergeSort\s*\(\s*numbers\s*\)\s*;?\s*console\.log\s*\(\s*['"]Sorted array:\s*['"]\s*\)\s*;?\s*console\.log\s*\(\s*numbers\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  if (array.length <= 1) {
    return;
  }

  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);

  mergeSort(leftPart);
  mergeSort(rightPart);

  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }
  
  while (rightArrayIndex < rightPart.length) {
    array[sortedIndex] = rightPart[rightArrayIndex];
    rightArrayIndex += 1;
    sortedIndex += 1;
  }

}
const numbers = [4, 10, 6, 14, 2, 1, 8, 5];
console.log('Unsorted array: ');
console.log(numbers);
mergeSort(numbers);
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function mergeSort(array) {
  if (array.length <= 1) {
    return;
  }

  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);

  mergeSort(leftPart);
  mergeSort(rightPart);

  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }

  while (rightArrayIndex < rightPart.length) {
    array[sortedIndex] = rightPart[rightArrayIndex];
    rightArrayIndex += 1;
    sortedIndex += 1;
  }
}

const numbers = [4, 10, 6, 14, 2, 1, 8, 5];
console.log('Unsorted array:');
console.log(numbers);
mergeSort(numbers);
console.log('Sorted array:');
console.log(numbers);
```
