---
id: 69b279bae0f20eda63163cf8
title: 14단계
challengeType: 1
dashedName: step-14
---

# --description--

이전 단계에서 만든 `if`와 `else` 문장은 정렬된 배열에 요소를 할당할 것입니다.

정렬된 배열에 할당된 각 요소는 목록에서 하나의 인덱스를 차지합니다. 따라서 각 할당 후에는 정렬된 배열에서 다음 인덱스로 이동해야 합니다.

`if`/`else` 블록 아래, 그러나 여전히 `while` 반복문 안에서 `sortedIndex`를 `1`만큼 증가시키세요.

# --hints--

`sortedIndex` 절 다음에 `1`의 현재 값을 `else`만큼 증가시켜야 합니다.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
