---
id: 69b2762b17d23fbee40a6c00
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

병합 정렬 알고리즘은 세 가지 작업을 수행합니다:

- 정렬되지 않은 수열을 하위 부분으로 나누기
- 하위 부분의 항목들을 정렬하기
- 정렬된 하위 부분을 병합하기

위 작업은 하위 부분들이 완전한 정렬된 수열로 병합될 때까지 재귀적으로 진행됩니다. 먼저 수열을 나누는 것부터 시작하세요.

먼저 `mergeSort` 함수 안에서 `middlePoint` 변수를 `const`로 선언하고 `Math.floor(array.length / 2)`을 할당하여 재귀 호출 전에 수열을 반으로 나눌 수 있게 하세요.

# --hints--

`middlePoint` 변수를 선언하고 `Math.floor(array.length / 2)`를 할당해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(const|let|var)\s+middlePoint\s*=\s*Math\.floor\(\s*array\.length\s*\/\s*2\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
