---
id: 69b279bae0f20eda63163cf2
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

다음으로, `leftPart`와 `rightPart`에 남아 있는 요소가 모두 있을 동안 계속되는 반복문을 만들어야 합니다.

이를 위해 두 가지 조건을 가진 `while` 반복문을 만드세요: 하나는 `leftArrayIndex`가 `leftPart`의 길이보다 작은지 확인하는 조건이고, 다른 하나는 `rightArrayIndex`가 `rightPart`의 길이보다 작은지 확인하는 조건입니다.

# --hints--

`while` 반복문을 만들어 `leftArrayIndex`가 `leftPart`의 길이보다 작은지, 그리고 `rightArrayIndex`이 `rightPart`의 길이보다 작은지 확인해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(code),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
