---
id: 69b278519ef46b2fd8bf1546
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

이미 `middlePoint` 값을 가지고 있습니다. 이를 사용해 배열을 두 부분으로 나누고 왼쪽 부분을 정렬할 준비를 하세요.

슬라이스 구문을 사용해 `array`의 왼쪽 절반을 추출하고 `leftPart`라는 변수에 할당하세요.

# --hints--

`leftPart` 함수 안에 `mergeSort` 변수를 선언하세요.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftPart\s*=/
);
```

`array.slice(0, middlePoint)`에 `leftPart`를 할당하세요.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /leftPart\s*=\s*array\.slice\(\s*0\s*,\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
