---
id: 69b279bae0f20eda63163cff
title: 21단계
challengeType: 1
dashedName: step-21
---

# --description--

함수 정의 외부에 다음 값을 가진 배열 `numbers`을 만드세요: `[4, 10, 6, 14, 2, 1, 8, 5]`.

# --hints--

`numbers`라는 이름의 배열을 다음 값들로 만드세요: `[4, 10, 6, 14, 2, 1, 8, 5]`

```js
const cleaned = __helpers.removeJSComments(code);
  assert.match(
    cleaned,
    /(const|let|var)\s+numbers\s*=\s*\[\s*4\s*,\s*10\s*,\s*6\s*,\s*14\s*,\s*2\s*,\s*1\s*,\s*8\s*,\s*5\s*\]\s*;?/
  );
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  if (array.length <= 1) {
    return;
  }

  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }
  
  while (rightArrayIndex < rightPart.length) {
    array[sortedIndex] = rightPart[rightArrayIndex];
    rightArrayIndex += 1;
    sortedIndex += 1;
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
