---
id: 69b279bae0f20eda63163cfe
title: 20단계
challengeType: 1
dashedName: step-20
---

# --description--

`mergeSort()` 함수를 테스트하기 전에, `array`의 길이가 `1` 이하일 때 함수 실행을 정지하는 기본 사례를 만들어야 합니다.

이 기본 사례가 재귀 호출을 정지시킵니다. 이것이 없으면, 리스트가 정렬되었거나 요소가 없을 때도 병합 정렬 작업이 계속 실행됩니다.

함수 선언 바로 다음에 `if`가 1 이하인지 확인하는 `array.length` 문을 만드세요. `if` 문 본문 안에 함수 실행을 정지하는 `return` 문을 추가하세요.

# --hints--

조건 `if`을 가진 `array.length <= 1` 문을 만들어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{/
);
```

`return` 문 본문에 `if` 문이 있어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{[\s\S]*return\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }
  
  while (rightArrayIndex < rightPart.length) {
    array[sortedIndex] = rightPart[rightArrayIndex];
    rightArrayIndex += 1;
    sortedIndex += 1;
  }
}
```
