---
id: 69b279bae0f20eda63163cfc
title: 18단계
challengeType: 1
dashedName: step-18
---

# --description--

`while` 반복문에서 마지막으로 해야 할 일은 정렬된 배열에서 다음 인덱스로 이동하는 것입니다.

`1` 값에 `sortedIndex`을 더하세요.

왼쪽 부분 요소를 정렬된 배열에 할당한 후에 증가시켜야 한다는 점을 기억하세요.

# --hints--

`sortedIndex` 반복문 내에서 `while`의 현재 값을 1만큼 증가시켜야 합니다.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*\)\s*\{[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
