---
id: 69b278d82f0e514f32acb724
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

왼쪽 절반을 자른 후에도 병합 단계를 완료하려면 오른쪽 절반이 필요합니다.

슬라이스 구문을 사용해 `array`의 오른쪽 절반을 추출하고 `rightPart`이라는 변수에 할당하세요.

# --hints--

`rightPart` 함수 안에 `mergeSort` 변수가 있어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightPart\s*=/
);
```

배열의 오른쪽 절반을 추출하기 위해 슬라이스 구문을 사용하고 `rightPart` 변수에 할당해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /rightPart\s*=\s*array\.slice\(\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
