---
id: 69b279bae0f20eda63163cef
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

`array` 리스트를 두 개의 별도 리스트로 나눈 후, 각 리스트를 계속 나누어 모든 요소가 각각의 리스트에 단독으로 있게 하세요. 숫자 하나만 있는 리스트는 항상 정렬된 상태입니다.

이를 위해 함수 안에서 재귀적으로 `mergeSort`을 호출하고, 호출 시 `leftPart`를 인수로 전달하세요.

# --hints--

같은 `mergeSort(leftPart)` 함수 안에서 `mergeSort`을 호출해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*leftPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
