---
id: 69b2713ff6c21fed2d3ad88b
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

병합 정렬은 분할 정복 원리를 사용해 데이터 모음을 정렬하는 정렬 알고리즘입니다. 즉, 모음을 더 작은 하위 부분으로 "분할"하고, 하위 부분을 독립적으로 정렬하여 "정복"한 다음, 정렬된 하위 부분을 병합합니다.

`mergeSort`라는 이름의 함수를 선언하고 `array` 매개변수를 받는 것부터 시작하세요.

# --hints--

`mergeSort` 매개변수를 하나 받는 `array`라는 이름의 함수를 선언하세요.

```js
assert.isFunction(mergeSort);
const params = __helpers.getFunctionParams(mergeSort.toString());
assert.lengthOf(params, 1);
assert.deepEqual(params[0].name, "array");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
