---
id: 69b279bae0f20eda63163cfb
title: 17단계
challengeType: 1
dashedName: step-17
---

# --description--

`while` 반복문 안에서 `leftArrayIndex`의 값을 1만큼 증가시키세요.

# --hints--

왼쪽 부분 요소를 정렬된 배열에 할당한 후 `leftArrayIndex`의 현재 값을 `1`만큼 증가시켜야 합니다.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*\)\s*\{[\s\S]*leftArrayIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
