---
id: 69b279bae0f20eda63163cf6
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

이전 단계에서, 더 작은 요소였기 때문에 `leftPart`의 요소를 `array` 목록에 할당했습니다. 하지만 항상 그런 것은 아닙니다. 비교하는 경우에 따라 오른쪽 요소가 더 작을 수도 있습니다.

`else`의 요소가 `leftPart`의 요소보다 작지 않을 때 실행할 `rightPart` 절을 만드세요.

`else` 블록 안에서 `rightPart[rightArrayIndex]`을 `array[sortedIndex]`에 할당하세요.

# --hints--

`else` 블록을 만들고 그 안에 `array[sortedIndex] = rightPart[rightArrayIndex]`을 포함해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /else\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
}
```
