---
id: 69b279bae0f20eda63163cfd
title: 19단계
challengeType: 1
dashedName: step-19
---

# --description--

마지막으로, 오른쪽 절반에 남아 있는 값을 모두 `array`에 복사하세요.

이를 위해 `while`가 `rightArrayIndex`보다 작을 동안 실행되는 `rightPart.length` 반복문을 추가하고, `rightPart[rightArrayIndex]`를 `array[sortedIndex]`에 할당하며 `rightArrayIndex`와 `sortedIndex`을 모두 증가시키세요.

# --hints--

조건 `while`를 가진 `rightArrayIndex < rightPart.length` 반복문을 만드세요.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

반복문 본문에서는 `rightPart[rightArrayIndex]`을 `array[sortedIndex]`로 이동해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?/
);
```

할당 후 `rightArrayIndex`와 `sortedIndex`을 모두 증가시키세요.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*rightArrayIndex\s*(\+\+|\+=\s*1)\s*;?[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
