---
id: 69b279bae0f20eda63163cf4
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

`if` 조건이 `true`로 평가되면, `leftPart` 배열의 요소가 `rightPart` 배열에서 비교 대상인 요소보다 작다는 뜻입니다.

그럴 경우, 정렬된 배열 내 `leftArrayIndex`에 `leftPart`에서 찾은 요소를 넣을 수 있습니다.

`if` 블록 안에서 `leftPart[leftArrayIndex]`을 `array[sortedIndex]`에 할당하세요.

# --hints--

`leftPart[leftArrayIndex]` 안에서 `array[sortedIndex]`을 `if`에 할당해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /array\s*\[\s*sortedIndex\s*\]\s*=\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
