---
id: 69b279bae0f20eda63163cf5
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

`leftPart` 요소를 정렬된 배열에 할당한 후 `leftArrayIndex`을 `1`만큼 증가시키세요.

# --hints--

정렬된 배열에 요소를 삽입한 후 `leftArrayIndex`의 현재 값에 더해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /leftArrayIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
