---
id: 69b279bae0f20eda63163cf0
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

함수 본문 맨 아래에서 `mergeSort()` 함수를 다시 호출하세요. 이번에는 함수 호출에 `rightPart`을 인수로 전달하세요.

# --hints--

함수 본문 맨 아래에서 `mergeSort()`를 인수 `rightPart`와 함께 호출해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*rightPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
