---
id: 69b279bae0f20eda63163cf3
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

반복문 안에서 각 절반의 현재 요소를 비교하기 시작해서 더 작은 값을 먼저 병합하세요.

`if` 블록을 추가해 `leftPart`의 현재 요소가 `rightPart`의 현재 요소보다 작은지 확인하세요. 
# --hints--

`if` 조건은 `leftPart[leftArrayIndex]`가 `rightPart[rightArrayIndex]`보다 작은지 확인해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*<\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
