---
id: 69b279bae0f20eda63163cf1
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

이제 리스트 `leftPart`와 `rightPart`을 정렬하고 원래의 `array`에 병합할 차례입니다.

두 리스트의 요소를 비교하면서 더 작은 요소를 메인 리스트에 병합하면 됩니다. `leftPart`와 `rightPart`의 모든 인덱스에 대해 이 비교를 수행할 것입니다.

세 개의 변수를 만드세요: `leftArrayIndex`, `rightArrayIndex`, `sortedIndex`를 만들고 각각 `0`로 초기화하세요. 이 변수들은 정렬 과정에서 각 인덱스를 추적하는 데 도움이 됩니다.

# --hints--

`leftArrayIndex` 변수를 선언하고 `0`으로 초기화하세요.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftArrayIndex\s*=\s*0\s*;?/
);
```

`rightArrayIndex` 변수를 선언하고 `0`로 초기화하세요.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightArrayIndex\s*=\s*0\s*;?/
);
```

`sortedIndex` 변수를 선언하고 `0`로 초기화하세요.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+sortedIndex\s*=\s*0\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);

--fcc-editable-region--
  
--fcc-editable-region--
}
```
