---
id: 5d822fd413a79914d39e98cc
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

¡Bienvenido al proyecto Skyline de Variables CSS! Inicie agregando un elemento `link` que vincule su archivo `styles.css` dentro del elemento `head`.

# --hints--

Su código debe tener un elemento `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

No debe cambiar sus etiquetas `head` existentes. Asegúrese de no haber eliminado la etiqueta de cierre.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

Debe tener un elemento `link` vacío.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

Su elemento `link` debe estar dentro de su elemento `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Su elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Su elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
