---
id: 5d822fd413a79914d39e98cd
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

En CSS, puedes seleccionar todo con un asterisco. Agrega un borde a todo utilizando el selector `*`, y dándole un `border` de `1px solid black`. Este es un truco que ayuda a visualizar dónde están los elementos y su tamaño. Eliminarás esto más tarde.

# --hints--

Debes usar el selector `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Debes utilizar la propiedad `border` para estilizar todos los elementos.

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

Todos los elementos deben tener un borde de `1px solid black`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
