---
id: 5d822fd413a79914d39e98cf
title: Pasos 4
challengeType: 0
dashedName: step-4
---

# --description--

Puedes ver el `body` (es la caja más interna en tu página); la caja que lo rodea es el elemento `html`. Haz que tu `body` llene toda la ventana gráfica dándole una `height` de `100vh`. Elimina el `margin` predeterminado del `body` estableciendo el `margin` a `0`. Finalmente, configura la propiedad `overflow` en `hidden` para ocultar cualquier barra de desplazamiento que aparezca cuando algo se extienda más allá de la ventana gráfica.

# --hints--

Debes usar el selector `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Tu `body` debe tener una `height` de `100vh`.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

Tu `body` debe tener un `margin` de `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

Tu `body` debe tener la propiedad `overflow` configurada en `hidden`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--

```

