---
id: 6735a7370e0ae93a4577c689
title: Tengeneza mpangilio wa gazeti
challengeType: 25
dashedName: design-a-newspaper-layout
demoType: onClick
---

# --description--

Katika maabara hii, utatengeneza mpangilio wa gazeti ukitumia CSS Grid, ikijumuisha dhana kama safu za gridi na nguzo za gridi.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Ukurasa wako unapaswa kuwa na kipengee cha `main` chenye darasa la vitu la `newspaper-layout`.
1. `.newspaper-layout` inapaswa kujumuisha `header` yenye darasa la vitu la `title` inayobeba `h1` kuonyesha jina la gazeti.
1. `.newspaper-layout` inapaswa kujumuisha `article` yenye darasa la vitu la `feature-article` kwa ajili ya makala kuu ya habari.
1. `.feature-article` inapaswa kujumuisha kipengee cha `h2` kwa kichwa cha makala ikifuatiwa na kipengee cha `p` kwa maudhui ya makala.
1. Unapaswa kuongeza vipengee vitatu vya `article` kwa makala ndogo ndogo, vyenye madarasa ya vitu ya `small-article1`, `small-article2`, na `small-article3`.
1. Kila makala ndogo inapaswa kujumuisha kipengee cha `h3` kwa kichwa cha makala ikifuatiwa na kipengee cha `p` kwa maudhui ya makala.
1. `.newspaper-layout` inapaswa kujumuisha kipengee cha `article` chenye darasa la vitu la `secondary-article` kwa sehemu ya ziada ya habari.
1. `.secondary-article` inapaswa kujumuisha kipengee cha `h2` kwa kichwa cha makala ikifuatiwa na kipengee cha `p` kwa maudhui ya makala.
1. `.newspaper-layout` inapaswa kujumuisha `figure` yenye darasa la vitu la `cover-image` kuonyesha picha inayowakilisha maudhui ya gazeti.
1. Vipengee vyenye madarasa ya vitu ya `title`, `feature-article`, `secondary-article`, `cover-image`, `small-article1`, `small-article2`, na `small-article3` vinapaswa kuwa na sifa ya `grid-area` iliyowekwa kwa jina la darasa la vitu lile lile.
1. `.newspaper-layout` inapaswa kutumia CSS Grid na sifa ya `grid-template-areas` kuainisha mpangilio wa sehemu:
   - `.title` inapaswa kuenea katika safu ya juu.
   - `.feature-article` inapaswa kuenea nguzo mbili katika safu ya pili, na `.cover-image` katika nguzo ya tatu.
   - `.secondary-article` inapaswa kuenea nguzo mbili katika safu ya tatu, na `.cover-image` katika nguzo ya tatu.
   - Makala ndogo ndogo tatu zinapaswa kuonekana katika safu ya nne.
1. `.newspaper-layout` inapaswa kuwa na sifa ya `grid-template-columns` inayogawanya nafasi katika nguzo tatu sawa.
1. Unapaswa kuweka sifa ya `.newspaper-layout` ya `grid-template-rows` kuwa `auto` kwa safu ya kwanza na kugawanya nafasi iliyobaki kwa sehemu sawa kwa safu zingine.
1. Unapaswa kuongeza pengo kati ya vipande vya gridi.
1. Hakikisha picha ndani ya `.cover-image` inafaa ndani ya nafasi iliyotengwa kwa kuweka sifa yake ya `max-width` kuwa `100%`.

**Kumbuka:** Hakikisha umeunganisha karatasi yako ya mitindo katika HTML yako na kutumia CSS yako.

# --hints--

Unapaswa kuwa na kipengee cha `main`.

```js
assert.exists(document.querySelector('main'));
```

Kipengee chako cha `main` kinapaswa kuwa na sifa ya `class` ya `newspaper-layout`.

```js
assert.exists(document.querySelector('main.newspaper-layout'));
```

Unapaswa kuwa na kipengee cha `header` chenye sifa ya `class` ya `title` ndani ya kipengee chako cha `.newspaper-layout`.

```js
const header = document.querySelector(".newspaper-layout > header");
assert.exists(header);
assert(header.classList.contains("title"));
```

Kipengee chako cha `header` kinapaswa kujumuisha kipengee cha `h1`.

```js
const h1 = document.querySelector("header h1");
assert.exists(h1);
```

Kipengee chako cha `.title` kinapaswa kuwa na maandishi.

```js
assert.isAbove(document.querySelector(".title").innerText.length, 0);
```

Unapaswa kuwa na kipengee cha `article` kwa kila makala chenye madarasa ya vitu ya `feature-article`, `small-article1`, `small-article2`, `small-article3`, na `secondary-article`, mtawalia.

```js
const featureArticle = document.querySelector('article.feature-article');
const smallArticle1 = document.querySelector('article.small-article1');
const smallArticle2 = document.querySelector('article.small-article2');
const smallArticle3 = document.querySelector('article.small-article3');
const secondaryArticle = document.querySelector('article.secondary-article');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(smallArticle3);
assert.exists(secondaryArticle);
```

`.feature-article` inapaswa kujumuisha kipengee cha `h2` ikifuatiwa na kipengee cha `p`.

```js
const h2 = document.querySelector(".feature-article h2");
const p = document.querySelector(".feature-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

`.small-article1`, `.small-article2`, na `.small-article3` zinapaswa kujumuisha kipengee cha `h3` ikifuatiwa na kipengee cha `p`.

```js
for (let i = 1; i < 4; i++) {
    let h3 = document.querySelector(`.small-article${i} h3`);
    let p = document.querySelector(`.small-article${i} h3 + p`);
    assert.exists(h3);
    assert.exists(p);
}
```

`.secondary-article` inapaswa kujumuisha kipengee cha `h2` ikifuatiwa na kipengee cha `p`.

```js
const h2 = document.querySelector(".secondary-article h2");
const p = document.querySelector(".secondary-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

Unapaswa kuwa na kipengee cha `figure` chenye darasa la vitu la `cover-image` ndani ya `.newspaper-layout`.

```js
assert.exists(document.querySelector(".newspaper-layout figure.cover-image"));
```

`.cover-image` inapaswa kujumuisha kipengee cha `img`.

```js
assert.exists(document.querySelector(".cover-image img"));
```

Sifa ya `display` ya kipengee cha `.newspaper-layout` inapaswa kuwekwa kuwa `grid`.

```js
assert.equal(getComputedStyle(document.querySelector('.newspaper-layout')).display, 'grid');
```

Kipengee chako cha `.newspaper-layout` kinapaswa kuwa na sifa ya `grid-template-areas` yenye `title` inayojitenga katika nguzo zote tatu za safu ya kwanza.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
assert.match(newspaperLayout.gridTemplateAreas, /^("|')title title title\1/);
```

Safu ya pili ya muundo wa gridi inapaswa kuwa na `feature-article` inayojitenga katika nguzo mbili na `cover-image`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromSecondRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").trimStart();
assert.match(fromSecondRow, /^("|')feature-article feature-article cover-image\1/);
```

Safu ya tatu ya muundo wa gridi inapaswa kuwa na `secondary-article` inayojitenga katika nguzo mbili na `cover-image`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromThirdRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromThirdRow, /^("|')secondary-article secondary-article cover-image\1/);
```

Safu ya nne ya muundo wa gridi inapaswa kuwa na `small-article1`, `small-article2`, na `small-article3`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromFourthRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromFourthRow, /^("|')small-article1 small-article2 small-article3\1$/);
```

Kichaguzi chako cha `.newspaper-layout` kinapaswa kutumia sifa ya `grid-template-columns` kugawanya nafasi katika sehemu tatu sawa. Kumbuka unaweza kutumia kitendakazi cha `repeat` au kuainisha kila sehemu ya sehemu kwa mkono.

```js
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
```

Kichaguzi chako cha `.newspaper-layout` kinapaswa kuwa na sifa ya `grid-template-rows` iliyowekwa kuwa `auto 1fr 1fr 1fr`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateRows, ['auto 1fr 1fr 1fr', 'auto repeat(3, 1fr)']);
```

Unapaswa kuongeza pengo kati ya vipande vya gridi.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gap);
```

`.title` inapaswa kuwa na sifa ya `grid-area` ya `title`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.title')?.gridArea, 'title');
```

`.feature-article` inapaswa kuwa na sifa ya `grid-area` ya `feature-article`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.feature-article')?.gridArea, 'feature-article');
```

`.cover-image` inapaswa kuwa na sifa ya `grid-area` ya `cover-image`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cover-image')?.gridArea, 'cover-image');
```

`.secondary-article` inapaswa kuwa na sifa ya `grid-area` ya `secondary-article`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.secondary-article')?.gridArea, 'secondary-article');
```

`.small-article1` inapaswa kuwa na sifa ya `grid-area` ya `small-article1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article1')?.gridArea, 'small-article1');
```

`.small-article2` inapaswa kuwa na sifa ya `grid-area` ya `small-article2`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article2')?.gridArea, 'small-article2');
```

`.small-article3` inapaswa kuwa na sifa ya `grid-area` ya `small-article3`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article3')?.gridArea, 'small-article3');
```

`img` ndani ya `.cover-image` inapaswa kuwa na sifa ya `max-width` ya `100%`.

```js
const coverImage = document.querySelector(".cover-image");
const img = document.querySelector(".cover-image img");
assert.equal(getComputedStyle(img).maxWidth, "100%");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Newspaper Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Layout</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main class="newspaper-layout">
        <header class="title">
            <h1>The Daily Local News</h1>
        </header>

        <article class="feature-article">
            <h2>Breaking News: Volcano Eruption Disrupts Tourism</h2>
            <p>Recently, a volcano erupted in a popular tourist destination. The eruption has caused widespread panic
                and has disrupted tourism in the area. The volcano has been spewing lava and ash for several days now,
                and authorities are urging residents and tourists to evacuate the area immediately. The eruption has
                also caused a number of flights to be cancelled, leaving many tourists stranded. The situation is still
                developing, and authorities are working to contain the eruption and ensure the safety of everyone in the
                area.</p>
        </article>

        <article class="small-article1">
            <h3>Sports: Local Team Wins Championship</h3>
            <p>Hockey fans are celebrating today as the local team has won the championship. The team, which has been on
                a winning streak all season, clinched the title in a thrilling final match. Fans took to the streets to
                celebrate the victory, waving flags and chanting the team's name.</p>
        </article>

        <article class="small-article2">
            <h3>Health: Tips for a Balanced Diet</h3>
            <p>A diet high in calories, sugar, and unhealthy fats can lead to a variety of health problems, including
                obesity, diabetes, and heart disease. To maintain a healthy weight and reduce your risk of chronic
                diseases, it's important to eat a balanced diet that includes a variety of nutrient-rich foods. Here are
                some tips for eating a balanced diet:</p>
        </article>

        <article class="small-article3">
            <h3>Travel: Top 10 Destinations for 2025</h3>
            <p>Traveling is one of the best ways to experience new cultures, meet new people, and see the world. If
                you're looking for inspiration for your next trip, here are the top 10 destinations for 2025:</p>
        </article>

        <article class="secondary-article">
            <h2>Technology: The Rise of AI</h2>
            <p>Artificial intelligence (AI) is changing the way we live and work. From self-driving cars to virtual
                assistants, AI is revolutionizing the world around us. But what exactly is AI, and how does it work? In
                this article, we'll explore the rise of AI and its impact on society.</p>
        </article>

        <figure class="cover-image">
            <img src="https://cdn.freecodecamp.org/curriculum/labs/volcano.jpg" alt="Cover Image">
        </figure>
    </main>
</body>

</html>
```

```css
body {
    font-family: 'Georgia', 'Times New Roman', Times, serif;
    padding: 0;
    margin: 0;
    background-color: #f0f0f0;
}

.newspaper-layout {
    display: grid;
    grid-template-areas:
        "title title title"
        "feature-article feature-article cover-image"
        "secondary-article secondary-article cover-image"
        "small-article1 small-article2 small-article3";
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: auto 1fr 1fr 1fr;
    gap: 15px;
    background-color: #fff;
    padding: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    box-sizing: border-box;
    text-align: justify;
}

.title {
    grid-area: title;
    background-color: #1B1B32;
    color: white;
    padding: 15px;
    text-align: center;
}

.title h1 {
    font-size: 2.5rem;
    line-height: 1.3;
    margin: 0;
}

.feature-article,
.secondary-article {
    background-color: #f9f9f9;
    padding: 15px;
    font-size: 1rem;
}

.feature-article {
    grid-area: feature-article;
}

.secondary-article {
    grid-area: secondary-article;
}

.small-article1,
.small-article2,
.small-article3 {
    background-color: #e9e9e9;
    padding: 15px;
    font-size: 0.9rem;
}

.small-article1 {
    grid-area: small-article1;
}

.small-article2 {
    grid-area: small-article2;
}

.small-article3 {
    grid-area: small-article3;
}

.cover-image {
    grid-area: cover-image;
    display: flex;
    justify-content: center;
    align-items: center;
    background-color: #ddd;
    padding: 15px;
    margin: 0;
}

.cover-image img {
    max-width: 100%;
    height: 350px;
    width: 350px;
    border-radius: 5px;
}
```
