---
id: 5d822fd413a79914d39e98cd
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在 CSS 中，你可以用星號標的所有元素。為所有元素添加邊框，方法是使用 `*` 選擇器，並給它一個 `border`，值為 `1px solid black`。這是一個幫助你視覺化元素位置和大小的技巧。你稍後會移除它。

# --hints--

你應該使用 `*` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

你應該使用 `border` 屬性來設定所有元素的樣式。

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

所有元素都應該有 `1px solid black` 邊框。

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
