---
id: 5d822fd413a79914d39e98ce
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

同時為所有元素添加 `box-sizing` 的值為 `border-box`。這樣你添加的邊框就不會增加元素的大小。

# --hints--

你應該使用 `box-sizing` 屬性。

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('box-sizing'));
```

你應該使用現有的 `*` 選擇器。

```js
// Two selectors create two CSSStyleRule objects
assert.lengthOf(new __helpers.CSSHelp(document).getStyleDeclarations('*'), 1);
```

所有元素的 `box-sizing` 都應該是 `border-box`。

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
* {
  border: 1px solid black;
}

--fcc-editable-region--

```

