---
id: 5d822fd413a79914d39e98cc
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

歡迎來到 CSS 變量天際線項目！ 首先在 `head` 元素中添加一個 `link` 元素鏈接你的 `styles.css` 文件。

# --hints--

你的代碼中應該有一個 `link` 元素。

```js
assert.isNotNull(document.querySelector('link'));
```

你不應更改現有的 `head` 標籤。 確保你沒有刪除結束標籤。

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

你應該有一個 `link` 空元素。

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

`link` 元素應該在 `head` 元素中。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

你的 `link` 元素應具有一個 `rel` 屬性，其值爲 `stylesheet`。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

你的 `link` 元素應具有一個 `href` 屬性，其值爲 `styles.css`。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
