---
id: 5d822fd413a79914d39e98cf
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

你可以看到 `body`（它是頁面中最內層的框）；包圍它的框是 `html` 元素。透過給 `body` 設定 `height` 為 `100vh`，讓它填滿整個視窗。透過將 `margin` 設為 `0`，移除 `body` 的預設 `margin`。最後，將 `overflow` 屬性設為 `hidden`，以隱藏當內容超出視窗時出現的任何捲軸。

# --hints--

你應該使用 `body` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

你的 `body` 應該有 `100vh` 的 `height`。

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

你的 `body` 應該有 `0` 的 `margin`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

你的 `body` 應該將 `overflow` 屬性設為 `hidden`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--
    
```

