---
id: 5d822fd413a79914d39e98cc
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

歡迎來到 CSS Variables Skyline 專案！首先為你的 `head` 元素添加一個 `link` 元素，將你的 `styles.css` 檔案連結進去。

# --hints--

你的程式碼應該有一個 `link` 元素。

```js
assert.isNotNull(document.querySelector('link'));
```

你不應該更改你現有的 `head` 頁籤。請確保你沒有刪除結束標籤。

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

你應該有一個空的 `link` 元素。

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

你的 `link` 元素應該位於你的 `head` 元素內。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

你的 `link` 元素應該有一個 `rel` 屬性，其值為 `stylesheet`。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

你的 `link` 元素應該有一個 `href` 屬性，其值為 `styles.css`。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
