---
id: 5d822fd413a79914d39e98ed
title: 步驟 34
challengeType: 0
dashedName: step-34
---

# --description--

將 `.fb4` 的位置相對於目前位置移動，為它添加 `position` 的 `relative` 和 `left` 的 `10%`。對 `.fb5` 也做相同操作，但使用 `right` 取代 `left`。這將涵蓋建築物之間剩餘的白色空間。

# --hints--

你應該給 `.fb4` 一個 `position` 為 `relative`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb4')?.position, 'relative');
```

你應該給 `.fb4` 一個 `left` 為 `10%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb4')?.left, '10%');
```

你應該給 `.fb5` 一個 `position` 為 `relative`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb5')?.position, 'relative');
```

你應該給 `.fb5` 一個 `right` 為 `10%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb5')?.right, '10%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>

    <div class="foreground-buildings">
      <div></div>
      <div></div>
      <div class="fb1"></div>
      <div class="fb2"></div>
      <div></div>
      <div class="fb3"></div>
      <div class="fb4"></div>
      <div class="fb5"></div>
      <div class="fb6"></div>
      <div></div>
      <div></div>
    </div>
  </body>
</html>
```

```css
:root {
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
  --building-color4: #538cc6;
}

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings, .foreground-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
  position: absolute;
  top: 0;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2);
}

.bb3 {
  width: 10%;
  height: 55%;
  background-color: var(--building-color3);
}

.bb4 {
  width: 11%;
  height: 58%;
  background-color: var(--building-color4);
}

.fb1 {
  width: 10%;
  height: 60%;
  background-color: var(--building-color4);
}

.fb2 {
  width: 10%;
  height: 40%;
  background-color: var(--building-color3);
}

.fb3 {
  width: 10%;
  height: 35%;
  background-color: var(--building-color1);
}
--fcc-editable-region--
.fb4 {
  width: 8%;
  height: 45%;
  background-color: var(--building-color1);
}

.fb5 {
  width: 10%;
  height: 33%;
  background-color: var(--building-color2);
}
--fcc-editable-region--
.fb6 {
  width: 9%;
  height: 38%;
  background-color: var(--building-color3);
}
    
```

