---
id: 5d822fd413a79914d39e98e1
title: 第 22 步
challengeType: 0
dashedName: step-22
---

# --description--

那樣行不通，因為你在 `.bb1` 中宣告的變數不會傳遞到 `.bb2` 和 `.bb3` 同層元素。這就是 CSS 的運作方式。基於此，變數通常會宣告在 `:root` 選擇器中。這是 CSS 中最高層級的選擇器；將你的變數放在那裡會讓它們在任何地方都能使用。為你的樣式表頂端添加 `:root` 選擇器，並將所有變數宣告移到那裡。

# --hints--

你應該在樣式表的頂部宣告一個 `:root` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(':root'));
```

你應該在 `:root` 選擇器中定義 `--building-color1`，其值為 `#aa80ff`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(':root')?.getPropertyValue('--building-color1')?.trim(), '#aa80ff');
```

你應該在 `:root` 選擇器中定義 `--building-color2`，其值為 `#66cc99`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(':root')?.getPropertyValue('--building-color2')?.trim(), '#66cc99');
```

你應該在 `:root` 選擇器中定義 `--building-color3`，其值為 `#cc6699`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(':root')?.getPropertyValue('--building-color3')?.trim(), '#cc6699');
```

你應該從 `.bb1` 中移除訂製屬性變數。

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1'));
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color2'));
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color3'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2, green);
}

.bb3 {
  width: 10%;
  height: 55%;
  background-color: var(--building-color3, pink);
}

.bb4 {
  width: 11%;
  height: 58%;
}
    
```

