---
id: 5d822fd413a79914d39e98d0
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

在 `body` 中創建一個具有 `background-buildings` 類別的 `div` 元素。這將是一個建築群組的容器。

# --hints--

你應該創建一個 `div` 元素。

```js
assert.exists(document.querySelector('div'));
```

你的 `div` 元素應該位於 `body` 內。

```js
assert.equal(document.querySelector('div')?.parentElement?.localName, 'body');
```

你的 `div` 元素應該有一個 `background-buildings` 類別

```js
assert.isTrue([...document.querySelector('div')?.classList]?.includes('background-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

```

