---
id: 5895f70cf9fc0f352b528e65
title: Passport einrichten
challengeType: 2
forumTopicId: 301565
dashedName: set-up-passport
---

# --description--

Es ist Zeit, *Passport* einzurichten, damit Sie endlich beginnen können, einem Benutzer zu erlauben, sich zu registrieren oder in ein Konto einzuloggen. Zusätzlich zu Passport verwenden Sie Express-session, um Sitzungen zu verwalten. Express-session bietet viele erweiterte Funktionen, die Sie nutzen können, aber vorerst verwenden Sie nur die Grundlagen. Diese Middleware speichert die Sitzungs-ID als Cookie im Client und ermöglicht es uns, auf die Sitzungsdaten mit dieser ID auf dem Server zuzugreifen. So halten Sie persönliche Kontoinformationen aus dem Cookie heraus, das vom Client verwendet wird, um Ihrem Server mitzuteilen, dass Clients authentifiziert sind, und behalten den *Schlüssel*, um auf die auf dem Server gespeicherten Daten zuzugreifen.

`passport@~0.4.1` und `express-session@~1.17.1` sind bereits installiert und beide als Abhängigkeiten in Ihrer `package.json`-Datei aufgeführt.

Sie müssen die Sitzungseinstellungen konfigurieren und Passport initialisieren. Erstellen Sie zuerst die Variablen `session` und `passport`, um `express-session` bzw. `passport` zu require.

Dann richten Sie Ihre Express-App so ein, dass sie die Sitzung verwendet, indem Sie die folgenden Optionen definieren:

```javascript
app.use(session({
  secret: process.env.SESSION_SECRET,
  resave: true,
  saveUninitialized: true,
  cookie: { secure: false }
}));
```

Stellen Sie sicher, dass Sie `SESSION_SECRET` zu Ihrer `.env`-Datei hinzufügen und ihm einen zufälligen Wert geben. Dieser wird verwendet, um den Hash zu berechnen, mit dem Ihr Cookie verschlüsselt wird!

Nachdem Sie das alles gemacht haben, sagen Sie Ihrer Express-App, dass sie `passport.initialize()` und `passport.session()` **verwenden** soll.

Senden Sie Ihre Seite ab, wenn Sie denken, dass alles richtig ist. Wenn Sie auf Fehler stoßen, können Sie <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-passport-3" target="_blank" rel="noopener noreferrer nofollow">das bis zu diesem Punkt abgeschlossene Projekt ansehen</a>.

# --hints--

Passport und Express-session sollten Abhängigkeiten sein.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport',
    'Your project should list "passport" as a dependency'
  );
  assert.property(
    packJson.dependencies,
    'express-session',
    'Your project should list "express-session" as a dependency'
  );
```

Abhängigkeiten sollten korrekt requiret werden.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')passport("|')/,
    'You should have required passport'
  );
  assert.match(
    data,
    /require.*("|')express-session("|')/,
    'You should have required express-session'
  );
```

Die Express-App sollte die neuen Abhängigkeiten verwenden.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(data, /passport\.initialize/, 'Your express app should use "passport.initialize()"');
  assert.match(data, /passport\.session/, 'Your express app should use "passport.session()"');
```

Sitzung und Sitzungsschlüssel sollten korrekt eingerichtet sein.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /secret *:\s*process\.env(\.SESSION_SECRET|\[(?<q>"|')SESSION_SECRET\k<q>\])/,
    'Your express app should have express-session set up with your secret as process.env.SESSION_SECRET'
  );
```

