---
id: 5895f70cf9fc0f352b528e66
title: Serialisierung eines Benutzerobjekts
challengeType: 2
forumTopicId: 301563
dashedName: serialization-of-a-user-object
---

# --description--

Serialisierung und Deserialisierung sind wichtige Konzepte im Zusammenhang mit der Authentifizierung. Ein Objekt zu serialisieren bedeutet, seinen Inhalt in einen kleinen *Schlüssel* umzuwandeln, der dann wieder in das ursprüngliche Objekt deserialisiert werden kann. Dadurch können wir erkennen, wer mit dem Server kommuniziert hat, ohne bei jeder Anfrage für eine neue Seite die Authentifizierungsdaten wie Benutzername und Passwort senden zu müssen.

Um dies richtig einzurichten, benötigen Sie eine serialize-Funktion und eine deserialize-Funktion. In Passport können diese wie folgt erstellt werden:

```javascript
passport.serializeUser(cb);
passport.deserializeUser(cb);
```

Die an `serializeUser` übergebene Callback-Funktion wird mit zwei Argumenten aufgerufen: dem vollständigen Benutzerobjekt und einem Callback, das von Passport verwendet wird.

Das Callback erwartet zwei Argumente: einen Fehler, falls vorhanden, und einen eindeutigen Schlüssel, der den Benutzer identifiziert und im Callback zurückgegeben werden soll. Sie verwenden die `_id` des Benutzers im Objekt. Diese ist garantiert eindeutig, da sie von MongoDB generiert wird.

Ebenso wird `deserializeUser` mit zwei Argumenten aufgerufen: dem eindeutigen Schlüssel und einer Callback-Funktion.

Dieses Callback erwartet zwei Argumente: einen Fehler, falls vorhanden, und das vollständige Benutzerobjekt. Um das vollständige Benutzerobjekt zu erhalten, führen Sie eine Abfrage nach einer Mongo `_id` durch, wie unten gezeigt:


```javascript
passport.serializeUser((user, done) => {
  done(null, user._id);
});

passport.deserializeUser((id, done) => {
  myDataBase.findOne({ _id: new ObjectID(id) }, (err, doc) => {
    done(null, null);
  });
});
```

Fügen Sie die beiden oben genannten Funktionen zu Ihrem Server hinzu. Die `ObjectID`-Klasse stammt aus dem `mongodb`-Paket. `mongodb@~3.6.0` wurde bereits als Abhängigkeit hinzugefügt. Deklarieren Sie diese Klasse mit:

```javascript
const { ObjectID } = require('mongodb');
```

Der `deserializeUser` wirft einen Fehler, bis Sie die Datenbankverbindung eingerichtet haben. Kommentieren Sie daher vorerst den `myDatabase.findOne`-Aufruf aus und rufen Sie stattdessen `done(null, null)` in der `deserializeUser`-Callback-Funktion auf.

Senden Sie Ihre Seite ab, wenn Sie glauben, dass alles richtig ist. Wenn Sie auf Fehler stoßen, können Sie <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#serialization-of-a-user-object-4" target="_blank" rel="noopener noreferrer nofollow">das bis zu diesem Punkt abgeschlossene Projekt ansehen</a>.

# --hints--

Sie sollten das Benutzerobjekt korrekt serialisieren.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport.serializeUser/gi,
    'You should have created your passport.serializeUser function'
  );
  assert.match(
    data,
    /null,\s*user._id/gi,
    'There should be a callback in your serializeUser with (null, user._id)'
  );
```

Sie sollten das Benutzerobjekt korrekt deserialisieren.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport.deserializeUser/gi,
    'You should have created your passport.deserializeUser function'
  );
  assert.match(
    data,
    /null,\s*null/gi,
    'There should be a callback in your deserializeUser with (null, null) for now'
  );
```

MongoDB sollte als Abhängigkeit vorhanden sein.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'mongodb',
    'Your project should list "mongodb" as a dependency'
  );
```

Mongodb sollte korrekt importiert werden, einschließlich des ObjectId.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')mongodb\1/gi,
    'You should have required mongodb'
  );
  assert.match(
    data,
    /new ObjectID.*id/gi,
    'Even though the block is commented out, you should use new ObjectID(id) for when we add the database'
  );
```

