---
id: 5895f70df9fc0f352b528e6a
title: Neue Middleware erstellen
challengeType: 2
forumTopicId: 301551
dashedName: create-new-middleware
---

# --description--

So wie es ist, kann jeder Benutzer einfach zu `/profile` gehen, egal ob er authentifiziert ist oder nicht, indem er die URL eingibt. Sie möchten dies verhindern, indem Sie zuerst prüfen, ob der Benutzer authentifiziert ist, bevor Sie die Profilseite rendern. Dies ist das perfekte Beispiel dafür, wann man eine Middleware erstellen sollte.

Die Aufgabe besteht darin, die Middleware-Funktion `ensureAuthenticated(req, res, next)` zu erstellen, die überprüft, ob ein Benutzer authentifiziert ist, indem sie die Methode `isAuthenticated` von Passport auf dem `request` aufruft, welche prüft, ob `req.user` definiert ist. Wenn ja, sollte `next()` aufgerufen werden. Andernfalls können Sie einfach auf die Anfrage mit einer Umleitung zur Startseite zum Einloggen antworten.

Eine Implementierung dieser Middleware sieht so aus:

```javascript
function ensureAuthenticated(req, res, next) {
  if (req.isAuthenticated()) {
    return next();
  }
  res.redirect('/');
};
```

Erstellen Sie die oben genannte Middleware-Funktion und übergeben Sie dann `ensureAuthenticated` als Middleware an die Anfragen für die Profilseite vor dem Argument der GET-Anfrage:

```javascript
app
 .route('/profile')
 .get(ensureAuthenticated, (req,res) => {
    res.render('profile');
 });
```

Senden Sie Ihre Seite ab, wenn Sie denken, dass alles richtig ist. Wenn Sie auf Fehler stoßen, können Sie <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#create-new-middleware-8" target="_blank" rel="noopener noreferrer nofollow">das bis zu diesem Punkt abgeschlossene Projekt ansehen</a>.

# --hints--

Die Middleware `ensureAuthenticated` sollte implementiert und an die `/profile`-Route angehängt werden.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /ensureAuthenticated[^]*req.isAuthenticated/,
    'Your ensureAuthenticated middleware should be defined and utilize the req.isAuthenticated function'
  );
  assert.match(
    data,
    /profile[^]*get[^]*ensureAuthenticated/,
    'Your ensureAuthenticated middleware should be attached to the /profile route'
  );
```

Eine nicht authentifizierte GET-Anfrage an `/profile` sollte korrekt zu `/` umgeleitet werden.

```js
  const url = new URL("/profile", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Home page/,
    'An attempt to go to the profile at this point should redirect to the homepage since we are not logged in'
  );
```

