---
id: 5895f70ef9fc0f352b528e6b
title: Wie man ein Profil zusammenstellt
challengeType: 2
forumTopicId: 301554
dashedName: how-to-put-a-profile-together
---

# --description--

Jetzt, da Sie sicherstellen können, dass der Benutzer, der auf `/profile` zugreift, authentifiziert ist, können Sie die in `req.user` enthaltenen Informationen auf Ihrer Seite verwenden.

Übergeben Sie ein Objekt mit der Eigenschaft `username` und dem Wert `req.user.username` als zweiten Argument für die `render`-Methode der Profilansicht.

Gehen Sie dann zu Ihrer `profile.pug`-Ansicht und fügen Sie die folgende Zeile unter dem vorhandenen `h1`-Element und auf derselben Einrückungsebene hinzu:

```pug
h2.center#welcome Welcome, #{username}!
```

Dies erstellt ein `h2`-Element mit der Klasse `center` und der ID `welcome`, das den Text `Welcome, ` gefolgt vom Benutzernamen enthält.

Fügen Sie außerdem in `profile.pug` einen Link ein, der auf die `/logout`-Route verweist, welche die Logik zum Deauthentifizieren eines Benutzers enthält:

```pug
a(href='/logout') Logout
```

Senden Sie Ihre Seite ab, wenn Sie denken, dass alles richtig ist. Wenn Sie auf Fehler stoßen, können Sie <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#how-to-put-a-profile-together-9" target="_blank" rel="noopener noreferrer nofollow">das bis zu diesem Punkt abgeschlossene Projekt ansehen</a>.

# --hints--

Sie sollten eine Pug-Render-Variable korrekt zu `/profile` hinzufügen.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /username:( |)req.user.username/,
    'You should be passing the variable username with req.user.username into the render function of the profile page'
  );
```

