---
id: 5895f70cf9fc0f352b528e67
title: Implementieren der Serialisierung eines Passport-Benutzers
challengeType: 2
forumTopicId: 301556
dashedName: implement-the-serialization-of-a-passport-user
---

# --description--

Sie laden kein tatsächliches Benutzerobjekt, da die Datenbank nicht eingerichtet ist. Stellen Sie beim Starten des Servers einmalig eine Verbindung zur Datenbank her und halten Sie diese Verbindung während des gesamten Lebenszyklus der App aufrecht. Fügen Sie dazu die Verbindungszeichenfolge Ihrer Datenbank (zum Beispiel: `mongodb+srv://<username>:<password>@cluster0-jvwxi.mongodb.net/?retryWrites=true&w=majority`) in die Umgebungsvariable `MONGO_URI` ein. Diese wird in der Datei `connection.js` verwendet.

*Wenn Sie Probleme haben, eine kostenlose Datenbank bei MongoDB Atlas einzurichten, sehen Sie sich dieses <a href="https://www.freecodecamp.org/news/get-started-with-mongodb-atlas/" target="_blank" rel="noopener noreferrer nofollow">Tutorial</a> an.*

Jetzt möchten Sie eine Verbindung zu Ihrer Datenbank herstellen und dann auf Anfragen hören. Der Zweck ist, keine Anfragen zuzulassen, bevor Ihre Datenbank verbunden ist oder wenn ein Datenbankfehler vorliegt. Um dies zu erreichen, fassen Sie Ihre Serialisierung und App-Routen in den folgenden Code ein:

```javascript
myDB(async client => {
  const myDataBase = await client.db('database').collection('users');

  // Be sure to change the title
  app.route('/').get((req, res) => {
    // Change the response to render the Pug template
    res.render('index', {
      title: 'Connected to Database',
      message: 'Please login'
    });
  });

  // Serialization and deserialization here...

  // Be sure to add this...
}).catch(e => {
  app.route('/').get((req, res) => {
    res.render('index', { title: e, message: 'Unable to connect to database' });
  });
});
// app.listen out here...
```

Stellen Sie sicher, dass Sie den `myDataBase`-Code in `deserializeUser` auskommentieren und bearbeiten Sie Ihre `done(null, null)`, um `doc` einzuschließen.

Senden Sie Ihre Seite ab, wenn Sie denken, dass alles richtig ist. Wenn Sie auf Fehler stoßen, können Sie <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implement-the-serialization-of-a-passport-user-5" target="_blank" rel="noopener noreferrer nofollow">das bis zu diesem Punkt abgeschlossene Projekt ansehen</a>.

# --hints--

Die Datenbankverbindung sollte vorhanden sein.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Connected to Database/gi,
    'You successfully connected to the database!'
  );
```

Die Deserialisierung sollte jetzt korrekt die Datenbank verwenden, und `done(null, null)` sollte mit `doc` aufgerufen werden.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /null,\s*doc/gi,
    'The callback in deserializeUser of (null, null) should be altered to (null, doc)'
  );
```

