---
id: 58965611f9fc0f352b528e6c
title: Benutzer abmelden
challengeType: 2
forumTopicId: 301560
dashedName: logging-a-user-out
---

# --description--

Die Logik zum Abmelden zu erstellen ist einfach. Die Route sollte den Benutzer einfach deautorisieren und zur Startseite umleiten, anstatt eine Ansicht zu rendern.

In Passport ist das Deautorisieren eines Benutzers so einfach wie das Aufrufen von `req.logout()` vor dem Umleiten. Fügen Sie diese `/logout`-Route hinzu, um das zu tun:

```js
app.route('/logout')
  .get((req, res) => {
    req.logout();
    res.redirect('/');
});
```

Vielleicht ist Ihnen aufgefallen, dass Sie fehlende Seiten (404) nicht behandeln. Die übliche Methode, dies in Node zu handhaben, ist mit folgender Middleware. Fügen Sie diese nach all Ihren anderen Routen hinzu:

```js
app.use((req, res, next) => {
  res.status(404)
    .type('text')
    .send('Not Found');
});
```

Reichen Sie Ihre Seite ein, wenn Sie denken, dass alles richtig ist. Wenn Sie auf Fehler stoßen, können Sie <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#logging-a-user-out-10" target="_blank" rel="noopener noreferrer nofollow">das bis zu diesem Punkt fertiggestellte Projekt ansehen</a>.

# --hints--

`req.logout()` sollte in Ihrer `/logout`-Route aufgerufen werden.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /req.logout/gi,
    'You should be calling req.logout() in your /logout route'
  );
```

`/logout` sollte zur Startseite umleiten.

```js
  const url = new URL("/logout", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Home page/gi,
    'When a user logs out they should be redirected to the homepage'
  );
```

