---
id: 58966a17f9fc0f352b528e6d
title: Registrierung neuer Benutzer
challengeType: 2
forumTopicId: 301561
dashedName: registration-of-new-users
---

# --description--

Jetzt müssen Sie einem neuen Benutzer auf Ihrer Website erlauben, ein Konto zu registrieren. Fügen Sie im `res.render` für die Startseite eine neue Variable zum übergebenen Objekt hinzu – `showRegistration: true`. Wenn Sie Ihre Seite aktualisieren, sollten Sie dann das Registrierungsformular sehen, das bereits in Ihrer `index.pug` Datei erstellt wurde. Dieses Formular ist so eingerichtet, dass es **POST** an `/register` sendet. Erstellen Sie also diese Route und lassen Sie sie das Benutzerobjekt in die Datenbank einfügen, indem Sie der folgenden Logik folgen.

Die Logik der Registrierungs-Route sollte wie folgt sein:

1. Registrieren Sie den neuen Benutzer
2. Authentifizieren Sie den neuen Benutzer
3. Leiten Sie zu `/profile` um

Die Logik von Schritt 1 sollte wie folgt sein:

1. Fragen Sie die Datenbank mit `findOne` ab
2. Wenn ein Fehler auftritt, rufen Sie `next` mit dem Fehler auf
3. Wenn ein Benutzer zurückgegeben wird, leiten Sie zurück zur Startseite um
4. Wenn kein Benutzer gefunden wird und keine Fehler auftreten, dann `insertOne` in die Datenbank mit dem Benutzernamen und Passwort. Solange dort keine Fehler auftreten, rufen Sie `next` auf, um zu Schritt 2 zu gelangen, der Authentifizierung des neuen Benutzers, für die Sie die Logik bereits in Ihrer `POST /login` Route geschrieben haben.

```js
app.route('/register')
  .post((req, res, next) => {
    myDataBase.findOne({ username: req.body.username }, (err, user) => {
      if (err) {
        next(err);
      } else if (user) {
        res.redirect('/');
      } else {
        myDataBase.insertOne({
          username: req.body.username,
          password: req.body.password
        },
          (err, doc) => {
            if (err) {
              res.redirect('/');
            } else {
              // The inserted document is held within
              // the ops property of the doc
              next(null, doc.ops[0]);
            }
          }
        )
      }
    })
  },
    passport.authenticate('local', { failureRedirect: '/' }),
    (req, res, next) => {
      res.redirect('/profile');
    }
  );
```

Senden Sie Ihre Seite ab, wenn Sie denken, dass alles richtig ist. Wenn Sie auf Fehler stoßen, können Sie <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#registration-of-new-users-11" target="_blank" rel="noopener noreferrer nofollow">das bis zu diesem Punkt abgeschlossene Projekt ansehen</a>.

**HINWEIS:** Ab diesem Punkt können Probleme im Zusammenhang mit der Verwendung des *Picture-in-Picture*-Browsers auftreten. Wenn Sie eine Online-IDE verwenden, die eine Vorschau der App innerhalb des Editors anbietet, wird empfohlen, diese Vorschau in einem neuen Tab zu öffnen.

# --hints--

Sie sollten eine `/register` Route haben und auf der Startseite ein Registrierungsformular anzeigen.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /showRegistration:( |)true/gi,
    'You should be passing the variable showRegistration as true to your render function for the homepage'
  );
  assert.match(
    data,
    /register[^]*post[^]*findOne[^]*username:( |)req.body.username/gi,
    'You should have a route that accepts a POST request on /register that queries the db with findOne and the query being username: req.body.username'
  );
```

Die Registrierung sollte funktionieren.

```js
  const url = code;
  const user = `freeCodeCampTester${Date.now()}`;
  const xhttp = new XMLHttpRequest();
  xhttp.onreadystatechange = function () {
    if (this.readyState == 4 && this.status == 200) {
      test(this);
    } else {
      throw new Error(`${this.status} ${this.statusText}`);
    }
  };
  xhttp.open('POST', url + '/register', true);
  xhttp.setRequestHeader('Content-type', 'application/x-www-form-urlencoded');
  xhttp.send(`username=${user}&password=${user}`);
  function test(xhttpRes) {
    const data = xhttpRes.responseText;
    assert.match(
      data,
      /Profile/gi,
      'Register should work, and redirect successfully to the profile.'
    );
  }
```

Das Einloggen sollte funktionieren.

```js
  const url = code;
  const user = `freeCodeCampTester${Date.now()}`;
  const xhttpReg = new XMLHttpRequest();
  xhttpReg.onreadystatechange = function () {
    if (this.readyState == 4 && this.status == 200) {
      login();
    } else {
      throw new Error(`${this.status} ${this.statusText}`);
    }
  };
  xhttpReg.open('POST', url + '/register', true);
  xhttpReg.setRequestHeader(
    'Content-type',
    'application/x-www-form-urlencoded'
  );
  xhttpReg.send(`username=${user}&password=${user}`);
  function login() {
    const xhttp = new XMLHttpRequest();
    xhttp.onreadystatechange = function () {
      if (this.readyState == 4 && this.status == 200) {
        test(this);
      } else {
        throw new Error(`${this.status} ${this.statusText}`);
      }
    };
    xhttp.open('POST', url + '/login', true);
    xhttp.setRequestHeader('Content-type', 'application/x-www-form-urlencoded');
    xhttp.send(`username=${user}&password=${user}`);
  }
  function test(xhttpRes) {
    const data = xhttpRes.responseText;
    assert.match(
      data,
      /Profile/gi,
      'Login should work if previous test was done successfully and redirect successfully to the profile.'
    );
    assert.match(
      data,
      new RegExp(user, 'g'),
      'The profile should properly display the welcome to the user logged in'
    );
  }
```

Das Ausloggen sollte funktionieren.

```js
  $.ajax({
    url: code + '/logout',
    type: 'GET',
    xhrFields: { withCredentials: true }
  }).then(
    (data) => {
      assert.match(data, /Home/gi, 'Logout should redirect to home');
    },
    (xhr) => {
      throw new Error(xhr.statusText);
    }
  );
```

Das Profil sollte nach dem Ausloggen nicht mehr funktionieren.

```js
  $.ajax({
    url: code + '/profile',
    type: 'GET',
    crossDomain: true,
    xhrFields: { withCredentials: true }
  }).then(
    (data) => {
      assert.match(
        data,
        /Home/gi,
        'Profile should redirect to home when we are logged out now again'
      );
    },
    (xhr) => {
      throw new Error(xhr.statusText);
    }
  );
```

