---
id: 589a8eb3f9fc0f352b528e72
title: Implementierung der sozialen Authentifizierung III
challengeType: 2
forumTopicId: 301558
dashedName: implementation-of-social-authentication-iii
---

# --description--

Der letzte Teil der Strategie besteht darin, das von GitHub zurückgegebene Profil zu verarbeiten. Wir müssen das Benutzerobjekt aus der Datenbank laden, falls es existiert, oder eines erstellen, falls nicht, und die Felder aus dem Profil füllen, dann das Benutzerobjekt zurückgeben. GitHub stellt uns eine eindeutige *id* innerhalb jedes Profils zur Verfügung, mit der wir suchen können, um den Benutzer zu serialisieren (bereits implementiert). Unten sehen Sie eine Beispielimplementierung, die Sie in Ihrem Projekt verwenden können – sie gehört in die Funktion, die das zweite Argument für die neue Strategie ist, direkt unterhalb von `console.log(profile);`:

```js
myDataBase.findOneAndUpdate(
  { id: profile.id },
  {
    $setOnInsert: {
      id: profile.id,
      username: profile.username,
      name: profile.displayName || 'John Doe',
      photo: profile.photos[0].value || '',
      email: Array.isArray(profile.emails)
        ? profile.emails[0].value
        : 'No public email',
      created_on: new Date(),
      provider: profile.provider || ''
    },
    $set: {
      last_login: new Date()
    },
    $inc: {
      login_count: 1
    }
  },
  { upsert: true, new: true },
  (err, doc) => {
    return cb(null, doc.value);
  }
);
```

`findOneAndUpdate` ermöglicht es Ihnen, ein Objekt zu suchen und zu aktualisieren. Falls das Objekt nicht existiert, wird es eingefügt und der Callback-Funktion zur Verfügung gestellt. In diesem Beispiel setzen wir immer `last_login`, erhöhen `login_count` um `1` und füllen die meisten Felder nur, wenn ein neues Objekt (neuer Benutzer) eingefügt wird. Beachten Sie die Verwendung von Standardwerten. Manchmal enthält ein zurückgegebenes Profil nicht alle Informationen oder der Benutzer hält sie privat. In diesem Fall behandeln Sie das, um einen Fehler zu vermeiden.

Sie sollten sich jetzt in Ihre App einloggen können. Probieren Sie es aus!

Senden Sie Ihre Seite ab, wenn Sie denken, dass alles richtig ist. Wenn Sie auf Fehler stoßen, können Sie <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implementation-of-social-authentication-iii-5" target="_blank" rel="noopener noreferrer nofollow">das bis zu diesem Punkt abgeschlossene Projekt ansehen</a>.

# --hints--

Die Einrichtung der GitHub-Strategie sollte abgeschlossen sein.

```js
  const url = new URL("/_api/auth.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /GitHubStrategy[^]*myDataBase/gi,
    'Strategy should use now use the database to search for the user'
  );
  assert.match(
    data,
    /GitHubStrategy[^]*cb/gi,
    'Strategy should return the callback function "cb"'
  );
```

