---
id: 589a69f5f9fc0f352b528e70
title: Implementierung der sozialen Authentifizierung
challengeType: 2
forumTopicId: 301559
dashedName: implementation-of-social-authentication
---

# --description--

Der grundlegende Ablauf dieser Art der Authentifizierung in Ihrer Anwendung ist:

1.  Der Benutzer klickt auf eine Schaltfläche oder einen Link, der ihn zu Ihrer Route führt, um sich mit einer bestimmten Strategie zu authentifizieren (z. B. GitHub).
2.  Ihre Route ruft `passport.authenticate('github')` auf, die ihn zu GitHub umleitet.
3.  Die Seite, auf der der Benutzer bei GitHub landet, ermöglicht ihm die Anmeldung, falls er noch nicht eingeloggt ist. Anschließend wird er gebeten, den Zugriff Ihrer Anwendung auf sein Profil zu genehmigen.
4.  Der Benutzer wird dann an Ihre Anwendung an eine bestimmte Callback-URL mit seinem Profil zurückgegeben, wenn er genehmigt wurde.
5.  Er ist nun authentifiziert, und Ihre Anwendung sollte prüfen, ob es sich um ein zurückkehrendes Profil handelt, oder es in Ihrer Datenbank speichern, falls nicht.

Strategien mit OAuth erfordern mindestens eine *Client-ID* und ein *Client-Geheimnis*, mit denen der Dienst überprüfen kann, von wem die Authentifizierungsanfrage stammt und ob sie gültig ist. Diese erhalten Sie von der Seite, mit der Sie die Authentifizierung implementieren möchten, wie GitHub, und sie sind einzigartig für Ihre Anwendung – **SIE DÜRFEN NICHT GETEILT WERDEN** und sollten niemals in einem öffentlichen Repository hochgeladen oder direkt im Code geschrieben werden. Üblich ist es, sie in Ihrer `.env`-Datei zu speichern und so zu referenzieren: `process.env.GITHUB_CLIENT_ID`. Für diese Aufgabe verwenden Sie die GitHub-Strategie.

Folgen Sie <a href="https://www.freecodecamp.org/news/how-to-set-up-a-github-oauth-application/" target="_blank" rel="noopener noreferrer nofollow">diesen Anweisungen</a>, um Ihre *Client-ID und Ihr Geheimnis* von GitHub zu erhalten. Setzen Sie die Homepage-URL auf Ihre Startseite (**nicht die URL des Projektcodes**) und die Callback-URL auf dieselbe Homepage-URL mit `/auth/github/callback` am Ende. Speichern Sie die Client-ID und Ihr Client-Geheimnis in der `.env`-Datei Ihres Projekts als `GITHUB_CLIENT_ID` und `GITHUB_CLIENT_SECRET`.

Fügen Sie in Ihrer `routes.js`-Datei `showSocialAuth: true` zur Homepage-Route hinzu, nach `showRegistration: true`. Erstellen Sie nun 2 Routen, die GET-Anfragen akzeptieren: `/auth/github` und `/auth/github/callback`. Die erste sollte nur passport aufrufen, um `'github'` zu authentifizieren. Die zweite sollte passport aufrufen, um `'github'` mit einer Fehlerumleitung zu `/` zu authentifizieren, und wenn das erfolgreich ist, zu `/profile` umleiten (ähnlich wie in Ihrem letzten Projekt).

Ein Beispiel, wie `/auth/github/callback` aussehen sollte, ist ähnlich wie bei der normalen Anmeldung:

```js
app.route('/login')
  .post(passport.authenticate('local', { failureRedirect: '/' }), (req,res) => {
    res.redirect('/profile');
  });
```

Reichen Sie Ihre Seite ein, wenn Sie glauben, dass alles richtig ist. Wenn Sie auf Fehler stoßen, können Sie <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implementation-of-social-authentication-3" target="_blank" rel="noopener noreferrer nofollow">das Projekt bis zu diesem Punkt ansehen</a>.

# --hints--

Route `/auth/github` sollte korrekt sein.

```js
  try {
    const res = await fetch(code + '/_api/routes.js');
    if (res.ok) {
      const data = await res.text();
      assert.match(
          data.replace(/\s/g, ''),
          /passport.authenticate.*?github/g,
          'Route auth/github should only call passport.authenticate with github'
        );
    } else {
      throw new Error(res.statusText);
    }
    const res2 = await fetch(code + '/_api/app-stack');
    if (res2.ok) {
      const data2 = JSON.parse(await res2.json());
      const dataLayer = data2.find(layer => layer?.route?.path === '/auth/github');
      assert.deepInclude(dataLayer?.route, { methods: {get: true}, path: "/auth/github"});
      assert.deepInclude(dataLayer?.route?.stack?.[0], {method: "get", name: "authenticate"});
    } else {
      throw new Error(res2.statusText);
    }
  } catch (err) {
    throw new Error(err);
  }
```

Route `/auth/github/callback` sollte korrekt sein.

```js
  try {
    const res = await fetch(code + '/_api/routes.js');
    if (res.ok) {
      const data = await res.text();
      assert.match(
        data.replace(/\s/g, ''),
        /failureRedirect:("|')\/\1/g,
        'Route auth/github/callback should accept a get request and call passport.authenticate for github with a failure redirect to home'
      );
    } else {
      throw new Error(res.statusText);
    }
    const res2 = await fetch(code + '/_api/app-stack');
    if (res2.ok) {
      const data2 = JSON.parse(await res2.json());
      const dataLayer = data2.find(layer => layer?.route?.path === '/auth/github/callback');
      assert.deepInclude(dataLayer?.route, { methods: {get: true}, path: "/auth/github/callback"});
      assert.deepInclude(dataLayer?.route?.stack?.[0], {method: "get", name: "authenticate"});
    } else {
      throw new Error(res2.statusText);
    }
  } catch (err) {
    throw new Error(err);
  }
```

