---
id: 589fc832f9fc0f352b528e79
title: Chatnachrichten senden und anzeigen
challengeType: 2
forumTopicId: 301562
dashedName: send-and-display-chat-messages
---

# --description--

Es ist Zeit, dass Sie Clients erlauben, eine Chatnachricht an den Server zu senden, damit diese an alle Clients gesendet wird! In Ihrer `client.js`-Datei sehen Sie bereits einen Codeblock, der behandelt, was passiert, wenn das Nachrichtenformular abgesendet wird.

```js
$('form').submit(function() {
  /*logic*/
});
```

Innerhalb des Codes zum Absenden des Formulars sollten Sie ein Ereignis auslösen, nachdem Sie `messageToSend` definiert haben, aber bevor Sie das Textfeld `#m` löschen. Das Ereignis sollte `'chat message'` heißen und die Daten sollten einfach `messageToSend` sein.

```js
socket.emit('chat message', messageToSend);
```

Auf Ihrem Server sollten Sie nun auf das Socket-Ereignis `'chat message'` hören, wobei die Daten `message` genannt werden. Sobald das Ereignis empfangen wird, sollte es das Ereignis `'chat message'` an alle Sockets mit `io.emit` senden und ein Datenobjekt mit `username` und `message` übermitteln.

In `client.js` sollten Sie jetzt auf das Ereignis `'chat message'` hören und, wenn es empfangen wird, ein Listenelement an `#messages` anhängen, das den Benutzernamen, einen Doppelpunkt und die Nachricht enthält!

An diesem Punkt sollte der Chat vollständig funktionsfähig sein und Nachrichten an alle Clients senden!

Senden Sie Ihre Seite ab, wenn Sie denken, dass alles richtig ist. Wenn Sie auf Fehler stoßen, können Sie <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#send-and-display-chat-messages-11" target="_blank" rel="noopener noreferrer nofollow">das bis zu diesem Punkt abgeschlossene Projekt ansehen</a>.

# --hints--

Der Server sollte auf `'chat message'` hören und es dann richtig senden.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")chat message('|")[^]*io.emit.*('|")chat message('|").*username.*message/s,
    'Your server should listen to the socket for "chat message" then emit to all users "chat message" with name and message in the data object'
  );
```

Der Client sollte die neuen Daten vom Ereignis `'chat message'` korrekt verarbeiten und anzeigen.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")chat message('|")[^]*messages.*li/s,
    'You should append a list item to #messages on your client within the "chat message" event listener to display the new message'
  );
```

