---
id: 589fc832f9fc0f352b528e79
title: Chatnachrichten senden und anzeigen
challengeType: 2
forumTopicId: 301562
dashedName: send-and-display-chat-messages
---

# --description--

It's time you start allowing clients to send a chat message to the server to emit to all the clients! In your `client.js` file, you should see there is already a block of code handling when the message form is submitted.

```js
$('form').submit(function() {
  /*logic*/
});
```

Innerhalb des Formularcodes solltest du nach Definition von `messageToSend`, aber vor Löschung des Textfelds `#m`, ein Ereignis emittieren. Das Ereignis sollte `'chat message'` heißen und die Daten sollten einfach `messageToSend` sein.

```js
socket.emit('chat message', messageToSend);
```

Auf deinem Server solltest du jetzt auf den Socket für das Ereignis `'chat message'` mit den als `message` benannten Daten warten. Sobald das Ereignis eingegangen ist, sollte es das Ereignis `'chat message'` an alle Buchsen, die `io.emit` verwenden, aussenden, wobei ein Datenobjekt gesendet wird, das den `username` und die `message` enthält.

Innerhalb `client.js` solltest du nun auf das Event `'chat message'` warten und bei Erhalt ein Listenelement, bestehend aus dem Benutzernamen, einem Doppelpunkt und der Nachricht, bei `#messages` hinzufügen!

Nun sollte der Chat voll funktionsfähig sein und Nachrichten an alle Clients verschicken!

Reiche deine Seite ein, wenn du davon ausgehst, alles richtig gemacht zu haben. Wenn du auf Fehler stößt, kannst du <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#send-and-display-chat-messages-11" target="_blank" rel="noopener noreferrer nofollow">das bis zu diesem Zeitpunkt abgeschlossene Projekt überprüfen</a>.

# --hints--

Server sollte auf das Ereignis `'chat message'` warten und dieses dann korrekt emittieren.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")chat message('|")[^]*io.emit.*('|")chat message('|").*username.*message/s,
    'Your server should listen to the socket for "chat message" then emit to all users "chat message" with name and message in the data object'
  );
```

Der Client sollte die neuen Daten aus dem Ereignis `'chat message'` richtig verarbeiten und anzeigen.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")chat message('|")[^]*messages.*li/s,
    'You should append a list item to #messages on your client within the "chat message" event listener to display the new message'
  );
```

