---
id: 589fc832f9fc0f352b528e78
title: Neue Benutzer ankündigen
challengeType: 2
forumTopicId: 301546
dashedName: announce-new-users
---

# --description--

Viele Chatrooms können ankündigen, wenn sich ein Benutzer verbindet oder trennt, und dies dann allen verbundenen Benutzern im Chat anzeigen. Da Sie bereits ein Ereignis beim Verbinden und Trennen senden, müssen Sie dieses Ereignis nur so anpassen, dass es eine solche Funktion unterstützt. Der logischste Weg ist, mit dem Ereignis drei Daten zu senden: den Benutzernamen des Benutzers, der sich verbunden/getrennt hat, die aktuelle Benutzeranzahl und ob sich dieser Benutzer verbunden oder getrennt hat.

Ändern Sie den Ereignisnamen zu `'user'` und übergeben Sie ein Objekt mit den Feldern `username`, `currentUsers` und `connected` (wobei `true` im Falle einer Verbindung oder `false` bei der Trennung des Benutzers gesendet wird). Achten Sie darauf, beide `'user count'`-Ereignisse zu ändern und das Trennungsereignis so einzustellen, dass es für das Feld `false` `connected` statt `true` wie beim Verbindungsereignis sendet.

```js
io.emit('user', {
  username: socket.request.user.username,
  currentUsers,
  connected: true
});
```

Nun hat Ihr Client alle notwendigen Informationen, um die aktuelle Benutzeranzahl korrekt anzuzeigen und anzukündigen, wenn sich ein Benutzer verbindet oder trennt! Um dieses Ereignis auf der Clientseite zu verarbeiten, sollten wir auf `'user'` hören, dann die aktuelle Benutzeranzahl aktualisieren, indem wir mit jQuery den Text von `#num-users` auf `'{NUMBER} users online'` ändern, sowie ein `<li>` an die ungeordnete Liste mit der ID `messages` mit `'{NAME} has {joined/left} the chat.'` anhängen.

Eine Implementierung könnte wie folgt aussehen:

```js
socket.on('user', data => {
  $('#num-users').text(data.currentUsers + ' users online');
  let message =
    data.username +
    (data.connected ? ' has joined the chat.' : ' has left the chat.');
  $('#messages').append($('<li>').html('<b>' + message + '</b>'));
});
```

Senden Sie Ihre Seite ab, wenn Sie denken, dass alles richtig ist. Wenn Sie auf Fehler stoßen, können Sie <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135/3#announce-new-users-10" target="_blank" rel="noopener noreferrer nofollow">das bis zu diesem Punkt abgeschlossene Projekt</a> ansehen.

# --hints--

Das Ereignis `'user'` sollte mit `username`, `currentUsers` und `connected` gesendet werden.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  // Regex is lenient to match both `username` and `name` as the key on purpose.
  assert.match(
    data,
    /io.emit.*('|")user\1.*name.*currentUsers.*connected/s,
    'You should have an event emitted named user sending name, currentUsers, and connected'
  );
```

Der Client sollte die neuen Daten vom Ereignis `'user'` korrekt verarbeiten und anzeigen.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")user\1[^]*num-users/s,
    'You should change the text of "#num-users" within on your client within the "user" event listener to show the current users connected'
  );
  assert.match(
    data,
    /socket.on.*('|")user\1[^]*messages.*li/s,
    'You should append a list item to "#messages" on your client within the "user" event listener to announce a user came or went'
  );
```

