---
id: 66e2edb23013ca596f8d9568
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Jetzt ist es Zeit, Ihre `getPunctuationCount`-Funktion zu testen.

Erstellen Sie eine `punctuationCount`-Variable und weisen Sie ihr das Ergebnis des Aufrufs der `getPunctuationCount`-Funktion mit dem Argument `"WHAT?!?!?!?!?"` zu.

Anschließend geben Sie Folgendes in der Konsole aus: `"Punctuation Count: [Punctuation count goes here]"`. Ersetzen Sie `[Punctuation count goes here]` durch den tatsächlichen Variablennamen. Sie können hier entweder Template-Strings oder String-Verkettung mit dem `+`-Operator verwenden.

# --hints--

Sie sollten eine `punctuationCount`-Variable erstellen.

```js
assert.isNotNull(punctuationCount)
```

Ihre `punctuationCount`-Variable sollte auf das Ergebnis von `getPunctuationCount("WHAT?!?!?!?!?")` gesetzt werden.

```js
assert.equal(punctuationCount, getPunctuationCount("WHAT?!?!?!?!?"));
```

Sie sollten Folgendes in der Konsole ausgeben: `"Punctuation Count: [Punctuation count goes here]"`. Ersetzen Sie `[Punctuation count goes here]` durch den tatsächlichen Variablennamen. Achten Sie darauf, hier die korrekte Syntax für String-Verkettung zu verwenden.

```js
assert.match(code, /console\.log\((?:('|"|`)Punctuation\s+Count:\s+('|"|`)\s+\+\s+punctuationCount|`Punctuation\s+Count:\s+\${punctuationCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```

