---
id: 66e2e99caa0cb3570f4822b4
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Jetzt ist es Zeit, Ihre `getConsonantCount`-Funktion zu testen.

Erstellen Sie eine `consonantCount`-Variable und weisen Sie ihr das Ergebnis des Aufrufs der `getConsonantCount`-Funktion mit dem Argument `"Coding is fun"` zu.

Loggen Sie danach Folgendes in die Konsole: `"Consonant Count: [Consonant count goes here]"`. Ersetzen Sie `[Consonant count goes here]` durch den tatsächlichen Variablennamen. Sie können hier Template-Strings oder String-Verkettung mit dem `+`-Operator verwenden.

# --hints--

Sie sollten eine `consonantCount`-Variable erstellen.

```js
assert.isNotNull(consonantCount)
```

Ihre `consonantCount`-Variable sollte auf das Ergebnis von `getConsonantCount("Coding is fun")` gesetzt werden.

```js
assert.equal(consonantCount, getConsonantCount("Coding is fun"));
```

Sie sollten die `consonantCount`-Variable mit dem Parameter `"Consonant Count: ${consonantCount}"` in der Konsole ausgeben.

```js
assert.match(code, /console\.log\((?:('|"|`)Consonant\s+Count:\s+('|"|`)\s+\+\s+consonantCount|`Consonant\s+Count:\s+\${consonantCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
