---
id: 66e2e5d890ac8a54c9d703bd
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Es ist Zeit, die Konsonanten zu zählen. Erstellen Sie eine `getConsonantCount`-Funktion mit einem `sentence`-Parameter.

Verwenden Sie innerhalb der Funktion eine Schleife, um die Anzahl der Konsonanten in dem `sentence` zu zählen, das beim Aufrufen der Funktion übergeben wird. Ein Konsonant ist jeder Buchstabe, der nicht eines der folgenden Zeichen ist: `"aeiou"`.

Ihre `getConsonantCount`-Funktion muss eine Zahl zurückgeben.

# --hints--

Sie sollten eine `getConsonantCount`-Funktion erstellen.

```js
assert.isFunction(getConsonantCount);
```

Ihre `getConsonantCount`-Funktion sollte einen `sentence`-Parameter haben.

```js
assert.match(getConsonantCount.toString(), /sentence/);
```

Ihre `getConsonantCount`-Funktion sollte eine Zahl zurückgeben.

```js
assert.isNumber(getConsonantCount("Coding is fun"))
```

Wenn der Satz `"Coding is fun"` ist, sollte die `getConsonantCount`-Funktion `7` zurückgeben.

```js
assert.strictEqual(getConsonantCount("Coding is fun"), 7);
```

Wenn der Satz `"Hello, World!"` ist, sollte die `getConsonantCount`-Funktion `7` zurückgeben.

```js
assert.strictEqual(getConsonantCount("Hello, World!"), 7);
```

Ihre Konsonantenzählung sollte nicht zwischen Groß- und Kleinschreibung unterscheiden.

```js
assert.strictEqual(getConsonantCount("Apples are tasty fruits"), 13);
assert.strictEqual(getConsonantCount("freeCodeCamp is awesome"), 11);
```

Ihre `getConsonantCount`-Funktion sollte für jeden Satz die korrekte Anzahl an Konsonanten zurückgeben.

```js
assert.strictEqual(getConsonantCount("I went to the store"), 9);
assert.strictEqual(getConsonantCount("The quick brown fox jumps over the lazy dog"), 24);
assert.strictEqual(getConsonantCount("The cat in the hat"), 9);
assert.strictEqual(getConsonantCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
