---
id: 66e2f06b191e305c00574e4d
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Zählen Sie schließlich die Anzahl der Wörter, indem Sie eine `getWordCount`-Funktion mit einem `sentence`-Parameter erstellen. Die Funktion soll die Gesamtzahl der Wörter im übergebenen Satz zurückgeben, wenn sie aufgerufen wird.

# --hints--

Sie sollten eine `getWordCount`-Funktion erstellen.

```js
assert.isFunction(getWordCount);
```

Ihre `getWordCount`-Funktion sollte einen `sentence`-Parameter haben.

```js
assert.match(getWordCount.toString(), /sentence/);
```

Ihre `getWordCount`-Funktion sollte eine Zahl zurückgeben.

```js
assert.isNumber(getWordCount("I love freeCodeCamp"))
```

Wenn der Satz `"When are you gonna start learning to code?"` ist, sollte die `getWordCount`-Funktion `8` zurückgeben.

```js
assert.strictEqual(getWordCount("When are you gonna start learning to code?"), 8);
```

Wenn der Satz `"What's going on?"` ist, sollte die `getWordCount`-Funktion `3` zurückgeben.

```js
assert.strictEqual(getWordCount("What's going on?"), 3);
```

Ihre Wortzählung sollte nicht zwischen Groß- und Kleinschreibung unterscheiden.

```js
assert.strictEqual(getWordCount("freeCodeCamp offers free coding tutorials online"), 6);
assert.strictEqual(getWordCount("You can learn HTML, CSS, JavaScript, and more on freeCodeCamp"), 10);
```

Ihre `getWordCount`-Funktion sollte die korrekte Wortanzahl für jeden Satz zurückgeben.

```js
assert.strictEqual(getWordCount("freeCodeCamp has a great community of kind people"), 8);
assert.strictEqual(getWordCount("The freeCodeCamp curriculum is constantly updated"), 6);
assert.strictEqual(getWordCount("freeCodeCamp teaches both front-end and back-end development"), 7);
```

Ihre `getWordCount`-Funktion sollte die korrekte Wortanzahl für einen leeren String oder einen String nur mit Leerzeichen zurückgeben.

```js
assert.strictEqual(getWordCount(""), 0);
assert.strictEqual(getWordCount("   "), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
