---
id: 66e2f376df6f315ee81de81a
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Testen Sie Ihre `getWordCount`, indem Sie eine `wordCount`-Variable erstellen, die auf den Aufruf der `getWordCount`-Funktion mit dem Satz `"I love freeCodeCamp"` gesetzt wird.

Anschließend geben Sie Folgendes in der Konsole aus: `"Word Count: [Word count goes here]"`. Ersetzen Sie `[Word count goes here]` durch den tatsächlichen Variablennamen. Sie können hier Template-Strings oder String-Verkettung mit dem `+`-Operator verwenden.

Damit ist Ihr Satzanalyse-Projekt abgeschlossen!

# --hints--

Sie sollten eine `wordCount`-Variable erstellen.

```js
assert.isNotNull(wordCount)
```

Ihre `wordCount`-Variable sollte auf das Ergebnis von `getWordCount("I love freeCodeCamp")` gesetzt werden.

```js
assert.equal(wordCount, getWordCount("I love freeCodeCamp"));
```

Sie sollten die `wordCount`-Variable mit dem Parameter `"Word Count: ${wordCount}"` in der Konsole ausgeben.

```js
assert.match(code, /console\.log\((?:('|"|`)Word\s+Count:\s+('|"|`)\s+\+\s+wordCount|`Word\s+Count:\s+\${wordCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

function getWordCount(sentence) {
  if (sentence.trim() === '') {
    return 0;
  }
  
  const words = sentence.trim().split(' ');
  let count = 0;

  for (const word of words) {
    if (word !== '') {
      count++;
    }
  }

  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

function getWordCount(sentence) {
  if (sentence.trim() === '') {
    return 0;
  }
  
  const words = sentence.trim().split(' ');
  let count = 0;

  for (const word of words) {
    if (word !== '') {
      count++;
    }
  }

  return count;
}

const wordCount = getWordCount("I love freeCodeCamp");
console.log(`Word Count: ${wordCount}`);
```
