---
id: 66e2d680e129e1423116a541
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop bauen Sie einen Satzanalysator, der einen Satz entgegennimmt und die Anzahl der Wörter, Vokale, Konsonanten und Satzzeichen ermittelt.

Beginnen Sie damit, eine `getVowelCount`-Funktion mit einem Parameter namens `sentence` zu erstellen. Ihre Funktion soll die Gesamtzahl der Vokale in einem Satz zurückgeben.

# --hints--

Sie sollten eine `getVowelCount`-Funktion erstellen.

```js
assert.isFunction(getVowelCount);
```

Ihre `getVowelCount`-Funktion sollte einen `sentence`-Parameter haben.

```js
assert.match(getVowelCount.toString(), /sentence/);
```

Ihre `getVowelCount`-Funktion sollte eine Zahl zurückgeben.

```js
assert.isNumber(getVowelCount("Apples are tasty fruits"))
```

Wenn der Satz `"Apples are tasty fruits"` ist, sollte die `getVowelCount`-Funktion `7` zurückgeben.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
```

Wenn der Satz `"Hello, World!"` ist, sollte die Funktion `3` zurückgeben.

```js
assert.strictEqual(getVowelCount("Hello, World!"), 3);
```

Ihre Vokalzählung sollte nicht zwischen Groß- und Kleinschreibung unterscheiden.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
assert.strictEqual(getVowelCount("freeCodeCamp is awesome"), 10);
```

Ihre `getVowelCount`-Funktion sollte für jeden Satz die korrekte Anzahl an Vokalen zurückgeben.

```js
assert.strictEqual(getVowelCount("I went to the store"), 6);
assert.strictEqual(getVowelCount("The quick brown fox jumps over the lazy dog"), 11);
assert.strictEqual(getVowelCount("The cat in the hat"), 5);
assert.strictEqual(getVowelCount(""), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
