---
id: 66e2df8c5676fe51cead8ec8
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Jetzt ist es Zeit, Ihre `getVowelCount`-Funktion zu testen.

Erstellen Sie eine `vowelCount`-Variable und weisen Sie ihr das Ergebnis des Aufrufs der `getVowelCount`-Funktion mit dem Argument `"Apples are tasty fruits"` zu.

Anschließend geben Sie Folgendes in der Konsole aus: `"Vowel Count: [vowel count goes here]"`. Ersetzen Sie `[vowel count goes here]` durch den tatsächlichen Variablennamen. Sie können hier entweder Template-Strings oder String-Konkatenation mit dem `+`-Operator verwenden.

# --hints--

Sie sollten eine `vowelCount`-Variable erstellen.

```js
assert.isNotNull(vowelCount)
```

Ihre `vowelCount`-Variable sollte auf das Ergebnis von `getVowelCount("Apples are tasty fruits")` gesetzt werden.

```js
assert.equal(vowelCount, getVowelCount("Apples are tasty fruits"));
```

Sie sollten die `vowelCount`-Variable mit dem Parameter `"Vowel Count: ${vowelCount}"` in der Konsole ausgeben.

```js
assert.match(code, /console\.log\((?:('|"|`)Vowel\s+Count:\s+('|"|`)\s+\+\s+vowelCount|`Vowel\s+Count:\s+\${vowelCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
