---
id: 66e2eab8a5638f57b637b7cc
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Sie sollten jetzt die Anzahl der Satzzeichen zählen.

Erstellen Sie eine `getPunctuationCount`-Funktion mit einem `sentence`-Parameter.

Innerhalb der Funktion erstellen Sie eine Schleife, um die Anzahl der Satzzeichen in dem `sentence` zu zählen, das beim Aufrufen der Funktion übergeben wird. Für unsere Zwecke ist ein Satzzeichen jedes Zeichen, das weder ein Leerzeichen (`" "`) noch ein Buchstabe ist.

Ihre `getPunctuationCount`-Funktion muss eine Zahl zurückgeben.

# --hints--

Sie sollten eine `getPunctuationCount`-Funktion erstellen.

```js
assert.isFunction(getPunctuationCount);
```

Ihre `getPunctuationCount`-Funktion sollte einen `sentence`-Parameter haben.

```js
assert.match(getPunctuationCount.toString(), /sentence/);
```

Ihre `getPunctuationCount`-Funktion sollte eine Zahl zurückgeben.

```js
assert.isNumber(getPunctuationCount("Coding is fun!"))
```

Wenn der Satz `"What's going on here?"` ist, sollte die `getPunctuationCount`-Funktion `2` zurückgeben.

```js
assert.strictEqual(getPunctuationCount("What's going on here?"), 2);
```

Wenn der Satz `"What????!"` ist, sollte die `getPunctuationCount`-Funktion `5` zurückgeben.

```js
assert.strictEqual(getPunctuationCount("What????!"), 5);
```

Ihre `getPunctuationCount`-Funktion sollte für jeden Satz die korrekte Anzahl der Satzzeichen zurückgeben.

```js
assert.strictEqual(getPunctuationCount("Be quick, sign up! freeCodeCamp awaits, friend!!!"), 6);
assert.strictEqual(getPunctuationCount("Guess what? freeCodeCamp is launching a new cert soon!"), 2);
assert.strictEqual(getPunctuationCount("freeCodeCamp, again? It's incredible!"), 4);
assert.strictEqual(getPunctuationCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
