---
id: 655cd899f8de09431eabb40c
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika mradi huu, utajifunza kuhusu miundo ya data kwa kujenga algoriti ya merge sort. 

Hii ni algoriti ya kupanga inayotumia kanuni ya kugawanya-na-kushinda ili kupanga makusanyo ya data. Yaani, inagawa mkusanyo kuwa sehemu ndogo ndogo, na kushinda sehemu hizo kwa kuzipanga kwa kujitegemea, kisha kuunganisha sehemu zilizopangwa.

Tengeneza kitendakazi kinachoitwa `merge_sort`. Kitendakazi hiki kitashughulikia zoezi la kupanga orodha ya nambari. 

Tumia neno kuu la `pass` ndani ya mwili wa kitendakazi.

# --hints--

Unapaswa kutangaza kitendakazi kinachoitwa `merge_sort`. Usisahau kutumia neno kuu la `pass` ndani ya mwili wa kitendakazi.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(merge_sort)
`))})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
