---
id: 6564aee9c077774ea49c3faf
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Algoriti ya merge sort hufanya vitendo vitatu kuu:

- Gawanya mfululizo usiopangwa wa vitu katika sehemu ndogo ndogo
- Panga vitu katika sehemu ndogo ndogo
- Unganisha sehemu ndogo ndogo zilizopangwa

Haya yote hufanyika kwa mfululizo wa kurudia hadi sehemu ndogo ndogo ziumanishwe kuwa mfululizo mzima uliopangwa. Anza kwa kugawanya mfululizo.

Kwanza, badilisha neno la `pass` na variable `middle_point`. Tumia operator ya kugawanya kwa nambari kamili (`//`) kugawanya urefu wa orodha ya `array` kwa 2 na weka thamani hiyo kwenye variable mpya ya `middle_point`. Kumbuka kuingiza msimbo wako kwa mpangilio sahihi.

# --hints--

Haupaswi kuwa na `pass` ndani ya mwili wa kitendakazi.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.notMatch(function_body, /^\s{4}pass/m);
  }
})
```

Unapaswa kutangaza variable ya `middle_point` na kuwekea `len(array) // 2` thamani yake.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}middle_point\s*=\s*len\(\s*array\s*\)\s*\/\/\s*2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
