---
id: 6564a9fe51964c229d5b7f4c
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Utahitaji kigezo kinachoonyesha mkusanyiko wa data unaopaswa kupangwa. Unda kigezo kinachoitwa `array` katika kitendakazi cha `merge_sort`.

# --hints--

Kitendakazi chako cha `merge_sort` kinapaswa kuchukua kigezo kimoja tu: `array`.

```js
({ test: () => assert(runPython(`
    import inspect
    str(inspect.signature(merge_sort)) == '(array)'    
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort():
    pass
--fcc-editable-region--
```
