---
id: 656627b47bd2d2a4afbc945d
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Katika hatua iliyopita, ulipata nambari ya katikati. Unaweza kuitumia kugawanya `array` vipande viwili na kuweka thamani ya kila sehemu katika vigezo vipya.

Tumia sintaksia ya slice kutoa nusu ya kushoto ya `array` na weka thamani yake katika kigezo kinachoitwa `left_part`.


# --hints--

Unapaswa kuwa na kigezo kinachoitwa `left_part` katika kitendakazi chako cha `merge_sort`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*(?!=)/m);
  }
})
```

Unapaswa kuweka thamani ya `array[:middle_point]` katika kigezo cha `left_part`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*=\s*array\s*\[\s*:\s*middle_point\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
