---
id: 656639771fed09b5c0e8fe71
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Sasa baada ya kugawanya orodha ya `array` kuwa orodha mbili tofauti, utaendelea kugawanya kila orodha hadi kila kipengele kiwe katika orodha yake binafsi. Orodha yenye nambari moja tu huwa imepangwa tayari.

Ili kufanya hivyo, itajiita `merge_sort` kwa mzunguko ndani ya kitendakazi chako. Usipitishe hoja yoyote kwa sasa.

# --hints--

Unapaswa kuita `merge_sort` chini kabisa ya mwili wa kitendakazi chako.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
--fcc-editable-region--
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
