---
id: 656710d1e0ec62253426db24
title: Hatua 18
challengeType: 20
dashedName: step-18
---

# --description--

Mzunguko wa `while` uliouunda unalinganisha kipengele kimoja kutoka `left_part` na kingine kutoka `right_part`, kisha unaongeza kipengele kidogo zaidi kwenye orodha kuu ya `array`.

Utaendelea na operesheni hii hadi pasipokuwepo tena vipengele vya kulinganisha. Lakini `left_part` bado inaweza kuwa na vipengele vilivyobaki wakati `right_part` haina, na kinyume chake.

Unda mzunguko mwingine wa `while` kunakili vipengele vilivyobaki katika `left_part` ndani ya orodha ya `array`. Tumia `left_array_index < len(left_part)` kwa hali ya `while`.

# --hints--

Mzunguko wako wa `while` unapaswa kuwa na hali hii: `left_array_index < len(left_part)`. Usisahau kutumia neno kuu la `pass` ndani ya mwili wa mzunguko.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}while\s+left_array_index\s*<\s*len\(\s*left_part\s*\):/m);
  }
})
```

Unapaswa kuwa na neno kuu la `pass` ndani ya mwili wa mzunguko wako wa `while`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\):[^}]*\bpass\b/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
--fcc-editable-region--

--fcc-editable-region--
```
