---
id: 656702f8b4cbd8cbf0a433c6
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Ndani ya kitendakazi chako, tengeneza mzunguko wa `while` unaolinganishwa kipengee katika `left_part` na kipengee katika `right_part`, na kuunganisha kipengee kidogo zaidi kwenye orodha kuu ya `array`.

Tengeneza masharti mawili kwa mzunguko: moja linalokagua kama `left_array_index` ni ndogo kuliko urefu wa `left_part` na sharti lingine linalokagua kama `right_array_index` ni ndogo kuliko urefu wa `right_part`.

# --hints--

Mzunguko wako wa `while` unapaswa kuwa na masharti haya: `left_array_index < len(left_part) and right_array_index < len(right_part)`. Usisahau kutumia neno kuu la `pass` ndani ya mwili wa mzunguko.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):/);
  }
})
```

Tumia neno kuu la `pass` ndani ya mwili wa mzunguko

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):\s*pass/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--

--fcc-editable-region--
```
