---
id: 656707e11e671ed54c96f7ec
title: Hatua 13
challengeType: 20
dashedName: step-13
---

# --description--

Wakati hali ya `if` inapotathminiwa kuwa `True`, ina maana kipengee kilicho katika orodha ya `left_part` ni kidogo kuliko kipengee kinacholinganishwa nacho katika orodha ya `right_part`.

Katika hali hiyo, unaweza weka kielezo cha `left_part` katika safu ya data iliyopangwa.

Ndani ya blokki ya `if`, ondoa `pass` na weka thamani ya `left_part[left_array_index]` kwa `array[sorted_index]`.

# --hints--

Unapaswa kubadilisha `pass` na `array[sorted_index] = left_part[left_array_index]` ndani ya mwili wa tamko la `if`.

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_body().is_equivalent("array[sorted_index] = left_part[left_array_index]")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            pass

--fcc-editable-region--
```
