---
id: 6569b743630ee592a65a7e2f
title: Hatua 26
challengeType: 20
dashedName: step-26
---

# --description--

Kabla ya kujaribu kitendakazi cha `merge_sort()`, unahitaji kuunda kesi ya msingi inayosimamisha utekelezaji wa kitendakazi wakati urefu wa `array` ni sawa au chini ya `1`.

Kesi hii ya msingi itasimamisha wito wa mzunguko. Bila yake, operesheni ya kupanga kwa kutumia merge sort ingeendelea kuendeshwa hata kama orodha tayari imepangwa au haina kipengele chochote ndani yake.

Mara moja baada ya tamko la kitendakazi, unda tamko la `if` lenye sharti hili: `len(array) <= 1`. Tumia neno kuu la `pass` ndani ya mwili wa kitendakazi.

# --hints--

Unapaswa kuunda tamko la `if` ili kukagua kama `len(array) <= 1`. Usisahau kutumia neno kuu la `pass` ndani ya mwili wa tamko la `if`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_function("merge_sort").find_ifs()[0].find_conditions()[0].is_equivalent("len(array) <= 1")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):

--fcc-editable-region--
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1
```
