---
id: 6569b4e0bd29d17d4b53b16c
title: Hatua 22
challengeType: 20
dashedName: step-22
---

# --description--

Sasa, utaiga mantiki ile ile ya mzunguko wa `while` kwa `right_part`.

Tengeneza mzunguko wa `while` unaoendelea wakati `right_array_index` ni chini ya `len(right_part)`, na tumia neno kuu la `pass` ndani ya mwili wa mzunguko.

# --hints--

Mzunguko wako wa `while` unapaswa kuwa na sharti hili: `right_array_index < len(right_part)`. Usisahau kutumia neno kuu la `pass` ndani ya mwili wa mzunguko.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}while\s+right_array_index\s*<\s*len\(\s*right_part\s*\):/m);
  }
})
```

Unapaswa kuwa na neno kuu la `pass` ndani ya mwili wa mzunguko wako wa `while`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+right_array_index\s*<\s*len\(\s*right_part\s*\):[^}]*\bpass\b/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--

--fcc-editable-region--
```
