---
id: 6577320da0d4c2e594d418e2
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

Pitia `right_part` kama hoja kwa wito wa `merge_sort()` uliouongeza katika hatua ya mwisho.

# --hints--

Unapaswa kupitisha `right_part` kama hoja kwa wito lako la mwisho la `merge_sort()`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*right_part\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort()
--fcc-editable-region--
```
