---
id: 657f59751d5a5c9b069ae0f3
title: Hatua 28
challengeType: 20
dashedName: step-28
---

# --description--

Unaweza kutumia variable ya `__name__` kubaini kama script ya Python inaendeshwa kama programu kuu au kama inaletwa kama moduli (msimbo uliandikwa katika faili nyingine ya Python).

Ikiwa thamani ya `__name__` imewekwa kuwa `'__main__'`, inaashiria kuwa script ya sasa ni programu kuu, si moduli.

Katika mradi huu, utatumia script ya sasa kama programu kuu.

Tengeneza tamko la `if` linalokagua kama thamani ya `__name__` ni `'__main__'`.

Tumia neno kuu la `pass` ndani ya mwili wa tamko la `if`.

# --hints--

Tamko lako la `if` linapaswa kukagua kama `__name__` ina thamani ya `'__main__'`

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent("__name__ == '__main__'")
    `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1

--fcc-editable-region--

--fcc-editable-region--
```
