---
id: 698367b9d1c6914a22095aec
title: 버블 정렬 알고리즘 구현하기
challengeType: 26
dashedName: implement-the-bubble-sort-algorithm
---

# --description--

이 실습에서는 버블 정렬 알고리즘을 구현하세요. 이 알고리즘은 정렬되지 않은 배열의 처음부터 시작해 인접한 값을 비교하며 정렬되지 않은 값을 배열 끝으로 "버블처럼 떠오르게" 합니다. 배열이 완전히 정렬될 때까지 반복합니다. 인접한 항목을 비교해 순서가 잘못된 경우 서로 교환합니다. 교환이 더 이상 발생하지 않을 때까지 배열을 반복해서 순회하며 정렬을 완료합니다.

이 방법은 배열을 여러 번 반복해서 순회해야 하며, 평균 및 최악의 경우 시간 복잡도가 이차함수적입니다. 간단하지만 대부분 상황에서는 비실용적입니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. 정수 배열을 입력으로 받아 이 정수들을 오름차순으로 정렬한 배열을 반환하는 `bubbleSort` 함수를 작성하세요.

# --hints--

`bubbleSort`이라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(bubbleSort);
```

`bubbleSort`는 오름차순으로 정렬된 배열을 반환해야 합니다.

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    bubbleSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`bubbleSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])`는 순서만 바뀌고 나머지는 변경되지 않은 배열을 반환해야 합니다.

```js
assert.sameMembers(
  bubbleSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`bubbleSort`는 내장 `.sort()` 메서드를 사용하지 않아야 합니다.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    bubbleSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function bubbleSort(array) {
  for (let i = 0; i < array.length; i++) {
    let swapped = false;
    for (let j = 1; j < array.length; j++) {
      if (array[j - 1] > array[j]) {
        let temp = array[j-1];
        array[j-1] =  array[j];
        array[j] = temp;
        swapped = true;
      }
    }
    if (swapped === false) {
      break;
    }
  }
  return array;
}
```
