---
id: 5dc17dc8f86c76b9248c6eb5
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

주석은 브라우저 표시에는 영향을 주지 않고 메시지를 남길 수 있게 해줍니다. 또한 코드를 비활성화할 수도 있습니다. HTML에서 주석은 `<!--`로 시작해서 여러 줄의 텍스트를 포함할 수 있으며 `-->`로 끝납니다.

다음은 `TODO: Remove h1`가 포함된 주석 예시입니다:

```html
<!-- TODO: Remove h1 -->
```

`p` 요소 위에 다음 텍스트가 포함된 주석을 추가하세요:

`TODO: Add link to cat photos`

# --hints--

주석은 `<!--`로 시작해야 합니다. 주석 시작을 정의하는 문자 중 하나 이상이 빠졌습니다.

```js
assert.match(code, /<!--/);
```

주석은 `-->`로 끝나야 합니다. 주석 끝을 정의하는 문자 중 하나 이상이 빠졌습니다.

```js
assert.match(code, /-->/);
```

코드에 불필요한 주석 시작/끝 문자가 없어야 합니다. 브라우저에 `<!--` 또는 `-->`가 추가로 표시되고 있습니다.

```js
const noSpaces = code.replace(/\s/g, '');
assert.isBelow(noSpaces.match(/<!--/g)?.length, 2)
assert.isBelow(noSpaces.match(/-->/g)?.length, 2);
```

주석은 `p` 요소 위에 있어야 합니다. 순서가 잘못되었습니다.

```js
assert.match(
  code.replace(/\s/g, ''),
  /<!--(.*?)--><p>everyonelovescutecatsonline!<\/p>/i
);
```

주석에 `TODO: Add link to cat photos` 텍스트가 포함되어야 합니다.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
--fcc-editable-region--
    <p>Everyone loves cute cats online!</p>
  </body>
</html>
```
