---
id: 5dc174fcf86c76b9248c6eb2
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 제목, 문단, 목록 같은 기본 HTML 요소를 사용해 고양이 사진 앱을 계속 만들어 나가겠습니다.

워크숍을 시작하려면 `h1` 요소를 추가하고 `CatPhotoApp` 텍스트를 넣으세요.

# --hints--

코드에 `CatPhotoApp` 텍스트가 있어야 합니다. 철자를 확인해 보세요.

```js
assert.match(code, /catphotoapp/i);
```

`h1` 요소에는 여는 태그가 있어야 합니다. 여는 태그 구문은 다음과 같습니다: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

`h1` 태그는 소문자로 작성해야 합니다. 관례상 모든 HTML 태그는 소문자로 씁니다.

```js
assert.notMatch(code, /<\/?H1>/);
```

`h1` 요소에는 닫는 태그가 있어야 합니다. 닫는 태그 구문은 다음과 같습니다: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

`h1` 요소의 텍스트는 `CatPhotoApp` 여야 합니다. 텍스트를 빼먹었거나 오타가 있거나 `h1` 요소의 여는 태그와 닫는 태그 사이에 있지 않습니다.

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
