---
id: 5dc1798ff86c76b9248c6eb3
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`h1` 요소 아래에 다음 텍스트가 있는 `h2` 요소를 추가하세요:

`Cat Photos`

# --hints--

`h1` 요소에는 여는 태그가 있어야 합니다. 여는 태그의 구문은 다음과 같습니다: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

`h1` 요소에는 닫는 태그가 있어야 합니다. 닫는 태그의 구문은 다음과 같습니다: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

`h1` 요소는 하나만 있어야 합니다. 추가된 요소를 제거하세요.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

`h1` 요소의 텍스트는 'CatPhotoApp'이어야 합니다. 텍스트가 빠졌거나 오타가 있습니다.

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```

`h2` 태그는 소문자여야 합니다. 관례상 모든 HTML 태그는 소문자로 작성합니다.

```js
assert.notMatch(code, /<\/?H2>/);
```

`h2` 요소에는 여는 태그가 있어야 합니다. 여는 태그의 구문은 다음과 같습니다: `<elementName>`.

```js
assert.exists(document.querySelector('h2'));
```

`h2` 요소에는 닫는 태그가 있어야 합니다. 닫는 태그는 `/`가 `<` 문자 바로 뒤에 옵니다.

```js
assert.match(code, /<\/h2\>/);
```

`h2` 요소의 텍스트는 `Cat Photos`이어야 합니다. 여는 `h2` 태그와 닫는 `h2` 태그 사이에는 `Cat Photos` 텍스트만 넣으세요.

```js
assert.equal(document.querySelector('h2')?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat photos');
```

`h2` 요소는 `h1` 요소 아래에 있어야 합니다. `h1` 요소가 더 중요하므로 `h2` 요소 위에 있어야 합니다.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H1'), collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
