---
id: 5dc24073f86c76b9248c6ebb
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML <dfn>속성</dfn>은 요소의 동작을 제어하기 위해 요소의 여는 태그 안에서 사용하는 특별한 단어입니다. `src` 속성은 `img` 요소에서 이미지의 URL(이미지가 위치한 곳)을 지정합니다.

다음은 freeCodeCamp 로고를 가리키는 `img` 속성이 있는 `src` 요소의 예입니다:

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

기존 `img` 요소 안에 다음 URL을 가진 `src` 속성을 추가하세요:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

코드에 `img` 요소가 있어야 합니다. `img` 요소를 삭제했거나 `src` 속성 값에 따옴표를 감싸지 않았을 수 있습니다.

```js
assert.exists(document.querySelector('img'));
```

`img` 요소에 `src` 속성이 있어야 합니다. 속성을 빼먹었거나 오타가 있을 수 있습니다. 요소 이름과 속성 이름 사이에 공백이 있는지 확인하세요.

```js
assert.exists(document.querySelector('img')?.src);
```

`img` 요소의 `src` 속성은 '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'로 설정되어야 합니다. URL을 빼먹었거나 오타가 있을 수 있습니다. URL 대소문자를 정확히 맞추세요.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

`img` 요소의 `src`을 올바른 URL로 설정했더라도, 속성 값은 항상 따옴표로 감싸는 것이 좋습니다.

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

`img` 요소는 `>` 또는 `/>`로 닫혀야 합니다.

```js
assert.match(code, /<img\s+src\s*=\s*("|')https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```
