---
id: 5dc23991f86c76b9248c6eb8
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

이전 단계에서 `h1`, `h2`, 주석, 그리고 `p` 요소를 `main` 요소 안에 넣었습니다. 이것을 *중첩*이라고 합니다. 중첩된 요소는 중첩된 요소가 속한 요소보다 오른쪽으로 두 칸 더 들여쓰기를 해야 합니다. 이 들여쓰기를 인덴테이션이라고 하며, HTML을 더 읽기 쉽게 만드는 데 사용합니다.

다음은 중첩과 인덴테이션의 예입니다:

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

아래 코드에서 `h1` 요소, `h2` 요소, 그리고 주석은 `main` 요소보다 두 칸 더 들여쓰기 되어 있습니다. 키보드의 스페이스 바를 사용해 `p` 요소 앞에 두 칸을 더 추가하여 올바르게 들여쓰기되도록 하세요.

# --hints--

코드에 텍스트가 `h2`인 `Cat Photos` 요소가 있어야 합니다. 실수로 삭제했거나, 여는 태그와 닫는 태그가 모두 없거나, 텍스트가 변경되었을 수 있습니다.

```js
assert.exists(document.querySelector('h2'));
assert.match(code, /<\/h2\>/);
assert.equal(document.querySelector('h2')?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(),'cat photos')
```

예제에서 `ul` 또는 `li` 요소를 추가하지 마세요.

```js
assert.notExists(document.querySelector('ul'));
assert.notExists(document.querySelector('li'));
```

`h2` 요소가 있는 줄의 인덴테이션을 변경하지 마세요. 여는 태그는 줄 시작에서 6칸 떨어진 위치에서 시작해야 합니다. 원래 인덴테이션으로 복원하려면 단계를 다시 시작하세요.

```js
assert.match(code.toLowerCase(), /<\/h1\>\s*\n\s{6}<h2>/);
```

코드에 주석이 있어야 합니다. 이전 단계에서 주석을 삭제했습니다.

```js
assert.match(code, /<!--.*-->/);
```

주석의 텍스트는 `TODO: Add link to cat photos`이어야 합니다. 주석의 텍스트나 공백을 변경하지 마세요.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i);
```

주석 요소가 있는 줄의 인덴테이션을 변경하지 마세요. 여는 태그는 줄 시작에서 6칸 떨어진 위치에서 시작해야 합니다. 원래 인덴테이션으로 복원하려면 단계를 다시 시작하세요.

```js
assert.match(
  code
    .toLowerCase(),
    /<\/h2>\s*\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/
);
```

코드에 `p` 요소가 있어야 합니다. 이전 단계에서 `p` 요소를 삭제했습니다.

```js
assert.exists(document.querySelector('p'));
```

`p` 요소의 텍스트는 `Everyone loves cute cats online!`이어야 합니다. `p` 요소의 텍스트, 공백, 구두점은 변경하지 마세요.

```js
assert.match(
  document
    .querySelector('p')
    ?.innerText?.trim().toLowerCase(),
    /^everyone\s+loves\s+cute\s+cats\s+online!$/
);
```

여는 `p` 태그는 `h2` 요소 및 주석 요소와 인덴테이션이 일치해야 합니다. 여는 태그는 줄 시작에서 6칸 떨어진 위치에서 시작해야 합니다.

```js
assert.match(code.toLowerCase(), /-->\s*\n\s{6}<p>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
    <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
    </main>
  </body>
</html>
```
