---
id: 5dc2385ff86c76b9248c6eb7
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5에는 서로 다른 콘텐츠 영역을 식별하는 요소들이 있습니다. 이러한 요소들은 HTML을 더 읽기 쉽게 만들고 검색 엔진 최적화(SEO)와 접근성에 도움을 줍니다.

`main` 요소는 HTML 문서 본문의 주요 콘텐츠를 나타내는 데 사용됩니다. `main` 요소 안의 콘텐츠는 문서에 고유해야 하며 문서의 다른 부분에서 반복되어서는 안 됩니다.

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

이 페이지의 주요 섹션을 `<main>` 요소 앞에 `h1` 여는 태그를 추가하고, `</main>` 요소 뒤에 `p` 닫는 태그를 추가하여 식별하세요.

# --hints--

`main` 요소에는 여는 태그가 있어야 합니다. 여는 태그의 구문은 다음과 같습니다: `<elementName>`.

```js
assert.exists(document.querySelector('main'));
```

`main` 태그는 소문자여야 합니다. 관례상 모든 HTML 태그는 소문자로 작성합니다.

```js
const mainTagMatches = code.matchAll(/<\/?main>/gi);
for (const match of mainTagMatches) {
  const tag = match[0]
  assert.strictEqual(
    tag,
    tag.toLowerCase()
  );
}
```

`main` 요소에는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 옵니다.

```js
assert.match(code, /<\/main\>/);
```

`main` 요소의 여는 태그는 `body` 요소의 여는 태그 아래에 있어야 합니다. 순서가 잘못되었습니다.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

`main` 요소의 여는 태그는 `h1` 요소 위에 있어야 합니다. 순서가 잘못되었습니다.

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('MAIN'), collection.indexOf('H1'));
```

`main` 요소의 닫는 태그는 `p` 요소 아래에 있어야 합니다. 순서가 잘못되었습니다.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode?.contains(pNode));
assert.match(pNode?.textContent.toLowerCase(), /everyone loves cute cats online/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    
--fcc-editable-region--
  </body>
</html>
```
