---
id: 5ddbd81294d8ddc1510a8e56
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

링크의 텍스트는 앵커(`a`) 요소의 여는 태그와 닫는 태그 사이에 위치해야 합니다.

다음은 텍스트가 `click here to go to freeCodeCamp.org`인 링크 예시입니다:

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

앵커 요소에 앵커 텍스트 `cat photos`를 추가하세요. 이것이 링크의 텍스트가 됩니다.

# --hints--

앵커(`a`) 요소에는 여는 태그가 있어야 합니다. 여는 태그 구문은 다음과 같습니다: `<elementName>`.

```js
assert.exists(document.querySelector('a'));
```

앵커(`a`) 요소에는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 있습니다.

```js
assert.match(code, /<\/a\>/);
```

앵커(`a`) 요소의 텍스트는 `cat photos`여야 합니다. 링크 텍스트를 앵커(`a`) 요소의 여는 태그와 닫는 태그 사이에 넣으세요.

```js
assert.equal(
  document.querySelector('a')?.innerText?.trim().toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
