---
id: 5dfa37b9eacea3f48c6300b0
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

순서 없는 목록을 만들려면 `ul` 요소를 사용할 수 있습니다.

`h3` 텍스트가 있는 `Things cats love:` 요소 다음에 순서 없는 목록(`ul`) 요소를 추가하세요. 이 시점에서는 아무것도 표시되지 않는다는 점을 참고하세요.

# --hints--

`ul` 요소는 여는 태그가 있어야 합니다. 여는 태그의 구문은 다음과 같습니다: `<elementName>`.

```js
assert.exists(document.querySelector('ul'));
```

`ul` 요소는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 있습니다.

```js
assert.match(code, /<\/ul>/);
```

`ul` 요소는 두 번째 `section` 요소의 닫는 태그 위에 있어야 합니다.

```js
const secondSectionLastElemNode =
  document.querySelectorAll('main > section')[1]?.lastElementChild;
assert.equal(secondSectionLastElemNode?.nodeName, 'UL');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
