---
id: 5dfa2407b521be39a3de7be1
title: 15단계
challengeType: 0
dashedName: step-15
---

# --description--

링크를 새 탭에서 열려면 앵커(`target`) 요소에 `a` 속성을 사용할 수 있습니다.

`target` 속성은 연결된 문서를 어디에서 열지 지정합니다. `target="_blank"`는 연결된 문서를 새 탭이나 창에서 엽니다.

다음은 `a` 속성이 있는 `target` 요소의 기본 구문입니다:

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

링크가 새 탭에서 열리도록 `target` 앵커(`_blank`) 요소의 여는 태그에 `cat photos` 속성을 값 `a`로 추가하세요.

# --hints--

`p` 요소에는 텍스트 `a`가 있는 중첩된 앵커(`cat photos`) 요소가 있어야 합니다. 삭제했거나 오타가 있을 수 있습니다.

```js
const anchor = document.querySelectorAll('p > a')[1];
assert.equal(
    anchor?.innerText?.toLowerCase().trim().replace(/\s+/g, ' '), 'cat photos'
);
```

앵커(`a`) 요소에 `target` 속성이 없습니다. 여는 태그 이름 뒤에 공백이 있는지, 그리고 모든 속성 이름 앞에 공백이 있는지 확인하세요.

```js
assert.isTrue(document.querySelectorAll('a')[1]?.hasAttribute('target'));
```

`target` 속성의 값은 `_blank`이어야 합니다. 값을 빼먹었거나 오타가 있을 수 있습니다. 속성 값은 반드시 따옴표로 감싸야 한다는 점을 기억하세요.

```js
assert.equal(document.querySelectorAll('a')[1]?.getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
