---
id: 5dfb6a35eacea3f48c6300b4
title: Step 26
challengeType: 0
dashedName: step-26
---

# --description--

figure 캡션(`figcaption`) 요소는 `figure` 요소 안에 포함된 이미지를 설명하는 캡션을 추가하는 데 사용합니다.

다음은 `figcaption` 캡션이 있는 `A cute cat` 요소 예시입니다:

```html
<figure>
  <img src="image.jpg" alt="A description of the image">
  <figcaption>A cute cat</figcaption>
</figure>
```

`figure` 요소에 중첩된 이미지 다음에, 텍스트가 다음과 같이 설정된 `figcaption` 요소를 추가하세요:

`Cats love lasagna.`

# --hints--

`figcaption` 요소에는 여는 태그가 있어야 합니다. 여는 태그의 구문은 다음과 같습니다: `<elementName>`.

```js
assert.exists(document.querySelector('figcaption'));
```

`figcaption` 요소에는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 있습니다.

```js
assert.match(code, /<\/figcaption\>/);
```

`figcaption` 요소는 `figure` 요소 안에 중첩되어야 합니다.

```js
assert.exists(document.querySelector('figure > figcaption'));
```

라자냐 `img` 요소는 `figure` 요소 안에 중첩되어야 합니다.

```js
assert.exists(document.querySelector('figure > img'))
assert.equal(
    document.querySelector('figure > img')?.getAttribute('src').toLowerCase(),
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

`figcaption` 요소에 중첩된 `figure` 요소는 `img` 요소 아래에 있어야 합니다. 순서가 잘못되었습니다.

```js
assert.equal(
  document.querySelector('figcaption')?.previousElementSibling.nodeName, 'IMG'
);
```

`figcaption` 요소의 텍스트는 `Cats love lasagna.`이어야 합니다. 텍스트를 빼먹었거나 오타가 있습니다.

```js
assert.match(
  document.querySelector('figcaption')?.innerText?.trim(), /Cats love lasagna.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          
--fcc-editable-region--          
        </figure>
      </section>
    </main>
  </body>
</html>
```
