---
id: 5ef9b03c81a63668521804ea
title: Step 38
challengeType: 0
dashedName: step-38
---

# --description--

지금까지 페이지에 추가한 모든 내용이 `body` 요소 안에 있다는 점을 확인하세요. 페이지에 렌더링되어야 하는 모든 페이지 콘텐츠 요소는 `body` 요소 안에 들어가야 합니다. 하지만 다른 중요한 정보는 `head` 요소 안에 들어갑니다.

`head` 요소는 문서의 제목, 스타일시트 링크, 스크립트 같은 메타데이터를 포함하는 데 사용됩니다. 메타데이터는 페이지에 직접 표시되지 않는 페이지에 관한 정보입니다.

`head` 요소 위에 `body` 요소를 추가하세요.

# --hints--

`body` 요소를 삭제했거나 여는 태그나 닫는 태그가 누락되었습니다.

```js
assert.exists(document.querySelector('body'));
assert.match(code, /<\/body>/);
```

`head` 요소에는 여는 태그가 있어야 합니다. 여는 태그의 구문은 다음과 같습니다: `<elementName>`.

```js
assert.match(code, /\<head\>/);
```

`head` 요소에는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 있습니다.

```js
assert.match(code, /\<\/head\>/);
```

`head` 요소는 여는 `body` 요소 태그 위에 있어야 합니다. 다른 곳에 두셨습니다.

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<\/head\>\<body\>/);
```

# --seed--

## --seed-contents--

```html
<html>
--fcc-editable-region--
  
--fcc-editable-region--
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
