---
id: 5ef9b03c81a63668521804ec
title: Step 40
challengeType: 0
dashedName: step-40
---

# --description--

페이지의 전체 내용이 `html` 요소 안에 중첩되어 있는 것을 확인하세요. `html` 요소는 HTML 페이지의 루트 요소로, 페이지의 모든 내용을 감쌉니다.

또한 `lang` 요소에 `html` 속성을 추가하여 페이지의 언어를 지정할 수 있습니다.

페이지의 언어가 영어임을 지정하려면 여는 `lang` 태그에 `en` 속성을 값 `html` 로 추가하세요.

# --hints--

`html` 요소를 삭제했거나 여는 태그나 닫는 태그가 없을 수 있습니다.

```js
assert.match(code, /\<html.*?\>/);
assert.match(code, /\<\/html\>/);
```

`html` 요소에 `lang` 속성이 값 `en` 으로 있어야 합니다. 속성이나 값이 빠졌거나 오타가 있을 수 있습니다.

```js
const extraSpacesRemoved = code.replace(/\s+/g, " ");
assert.match(extraSpacesRemoved, /\<html\s+lang\s*\=\s*("?|'?)en\1\s*\>/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<html>
--fcc-editable-region--
  <head>
    <title>CatPhotoApp</title>
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
