---
id: 5ef9b03c81a63668521804d2
title: Step 29
challengeType: 0
dashedName: step-29
---

# --description--

순서 있는 목록(`ol`)의 코드는 순서 없는 목록과 비슷하지만, 순서 있는 목록의 항목들은 화면에 표시될 때 번호가 붙습니다.

`h3` 요소 아래에 다음 세 개의 항목을 가진 순서 있는 목록을 추가하세요:

`flea treatment`
`thunder`
`other cats`

# --hints--

`ol` 요소에는 여는 태그가 있어야 합니다. 여는 태그의 구문은 `<elementName>`입니다.

```js
assert.exists(document.querySelector('ol'));
```

`ol` 요소에는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 옵니다.

```js
assert.match(code, /<\/ol>/);
```

`ol` 요소는 두 번째 `section` 요소의 닫는 태그 위에 있어야 합니다. 순서가 잘못되었습니다.

```js
assert.equal(document.querySelectorAll('main > section')[1]?.lastElementChild.nodeName, 'OL');
```

세 개의 `li` 요소는 `ol` 요소 안에 중첩되어야 합니다.

```js
assert.lengthOf(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ), 3
);
```

세 개의 `li` 요소가 있어야 하며, 텍스트는 `flea treatment`, `thunder`, `other cats` 중 어떤 순서로든 상관없습니다.

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText?.trim().replace(/\s+/g, ' ').toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

`ol` 요소는 하나만 있어야 합니다.

```js 
assert.lengthOf([...document.querySelectorAll('ol')], 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```
