---
id: 5ef9b03c81a63668521804e9
title: Step 37
challengeType: 0
dashedName: step-37
---

# --description--

기존 `freeCodeCamp.org` 텍스트를 앵커(`a`) 요소로 감싸 링크로 만드세요. `href` 속성은 `https://www.freecodecamp.org`로 설정해야 합니다.

# --hints--

앵커(`a`) 요소는 `p` 안의 `footer` 요소 안에 중첩되어야 합니다. 앵커(`a`) 요소에 여는 태그와 닫는 태그를 반드시 추가하세요.

```js
assert.isNotEmpty(document.querySelectorAll('footer > p > a'));
```

앵커(`a`) 요소에는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 있습니다.

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert.exists(aElemClosingTags);
assert.lengthOf(aElemClosingTags, 4);
```

앵커(`a`) 요소에는 값이 `href`인 `https://www.freecodecamp.org` 속성이 있어야 합니다. 속성이나 값이 빠졌거나 오타가 있을 수 있습니다.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(nestedAnchor?.getAttribute('href').toLowerCase(), 'https://www.freecodecamp.org');
```

링크의 텍스트는 `freeCodeCamp.org`여야 합니다. 텍스트가 빠졌거나 오타가 있을 수 있습니다.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(
  nestedAnchor?.innerText?.trim().toLowerCase(),
    'freecodecamp.org'
);
```

앵커(`a`) 요소를 중첩한 후 브라우저에 보이는 유일한 `p` 요소 내용은 `No Copyright - freeCodeCamp.org`이어야 합니다. `p`과 중첩된 앵커 요소의 텍스트, 띄어쓰기, 구두점을 다시 확인하세요.

```js
const pText = document.querySelectorAll('footer > p')[0]?.innerText?.trim().toLowerCase().replace(/\s+/g, ' ');
assert.match(pText, /^no copyright - freecodecamp\.org$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```
