---
id: 5ef9b03c81a63668521804ee
title: Step 41
challengeType: 0
dashedName: step-41
---

# --description--

모든 페이지는 `<!DOCTYPE html>`로 시작해야 합니다. 이 특별한 문자열은 <dfn>선언</dfn>으로 알려져 있으며, 브라우저가 업계 표준을 준수하도록 합니다.

`<!DOCTYPE html>`은 브라우저에 이 문서가 최신 버전인 HTML5 문서임을 알려줍니다.

이 선언을 코드의 첫 번째 줄에 추가하세요.

# --hints--

코드는 `<!DOCTYPE html>` 선언으로 시작해야 합니다. 선언을 빼먹었거나, 오타가 있거나, 코드의 첫 번째 줄이 아닐 수 있습니다.

```js
assert.match(code, /<\s*!doctype\s+html\s*>/i);
```

`<!DOCTYPE html>`는 문서 맨 위에 위치해야 합니다.

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /^<\!doctypehtml><html/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
<html lang="en">
  <head>
    <title>CatPhotoApp</title>
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
