---
id: 5ef9b03c81a63668521804d4
title: Step 34
challengeType: 0
dashedName: step-34
---

# --description--

`strong` 요소는 어떤 텍스트가 매우 중요하거나 긴급함을 나타낼 때 사용합니다.

방금 추가한 `figcaption`에서 `hate`가 매우 중요함을 나타내기 위해 `strong` 요소로 감싸세요.

# --hints--

`strong` 요소는 여는 태그가 있어야 합니다. 여는 태그의 구문은 다음과 같습니다: `<elementName>`.

```js
assert.exists(document.querySelector('strong'));
```

`strong` 요소는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 옵니다.

```js
assert.match(code, /<\/strong\>/);
```

`strong` 요소는 텍스트 `hate`에서 단어 `Cats hate other cats.`를 감싸야 합니다. 텍스트를 빼먹었거나 오타가 있을 수 있습니다.

```js
assert.equal(
  document
    .querySelectorAll('figcaption')[1]
    ?.querySelector('strong')
    ?.innerText?.trim().toLowerCase(), 'hate'
);
```

`figcaption`의 텍스트는 `Cats hate other cats.`이어야 합니다. 오타가 없는지, `strong` 요소의 여는 태그와 닫는 태그 주변에 필요한 공백이 있는지 확인하세요.

```js
const secondFigCaption = document.querySelectorAll('figcaption')[1];
assert.match(
   secondFigCaption?.innerText
      .replace(/\s+/gi, ' ')
      .trim(),
      /cats hate other cats\.?/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
--fcc-editable-region--
          <figcaption>Cats hate other cats.</figcaption>  
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```
