---
id: 5f07be6ef7412fbad0c5626b
title: 18단계
challengeType: 0
dashedName: step-18
---

# --description--

새로운 콘텐츠를 추가하기 전에 `section` 요소를 사용해 고양이 사진 콘텐츠와 앞으로 추가할 콘텐츠를 구분하세요.

`section` 요소는 문서에서 챕터, 헤더, 푸터 또는 기타 섹션과 같은 구역을 정의하는 데 사용됩니다. 이는 SEO와 접근성에 도움이 되는 의미론적 요소입니다.

```html
<section>
  <h2>Section Title</h2>
  <p>Section content...</p>
</section>
```

`h2` 요소, 두 개의 `p` 요소, 그리고 앵커(`a`) 요소를 `section` 요소 안에 중첩하세요.

# --hints--

`section` 요소는 여는 태그가 있어야 합니다. 여는 태그의 구문은 다음과 같습니다: `<elementName>`.

```js
assert.exists(document.querySelector('section'));
```

`section` 요소는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 옵니다.

```js
assert.match(code, /<\/section\s*>/i);
```

전체 `section` 요소는 `main` 요소의 여는 태그와 닫는 태그 사이에 있어야 합니다.

```js
assert.equal(document.querySelector('section')?.parentNode.nodeName, 'MAIN');
```

기존 `h2`, `p` 요소와 앵커(`a`) 요소는 `section` 요소의 여는 태그와 닫는 태그 사이에 있어야 합니다.

```js
const sectionEl = document.querySelector('section');
assert.isNotNull(sectionEl);
const childrenOfSection = [...sectionEl.childNodes];

const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert.lengthOf(foundElements, 4)
```

`h1` 요소는 `section` 요소 안에 중첩되어서는 안 됩니다.

```js
const sectionEl = document.querySelector('section');
assert.isNotNull(sectionEl);
const childrenOfSection = [...sectionEl.childNodes];

const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      
      <h2>Cat Photos</h2>
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
