---
id: 671141feba228a35cefba82d
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

기존 텍스트 `cute cats`를 `https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`로 연결되는 앵커 요소로 만드세요.

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`

# --hints--

첫 번째 `p` 요소 안에 새로운 앵커 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('p')?.querySelector('a'));
```

앵커 텍스트는 `cute cats`이어야 합니다.

```js
assert.strictEqual(document.querySelector('a')?.innerText?.trim(), "cute cats");
```

새 앵커 요소의 `href` 속성은 `https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`여야 합니다.

```js
assert.strictEqual(document.querySelector('a')?.href, "https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg");
```

코드에 `cute cats` 구문이 한 번만 있어야 합니다.

```js
assert.strictEqual(code.indexOf("cute cats"),code.lastIndexOf("cute cats")); 
```

`p` 요소의 텍스트는 여전히 `Everyone loves cute cats online!`이어야 합니다.

```js
assert.strictEqual(document.querySelector('p')?.innerText?.trim(), "Everyone loves cute cats online!");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
