---
id: 671141f948cbab359e74cc93
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

`p`을 단락으로 만들기 위해 `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.` 태그를 추가하세요.

# --hints--

`p` 주위에 `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.` 요소를 추가해야 합니다.

```js
const P = document.querySelectorAll('p')[1];
assert.exists(P);
assert.strictEqual(P.innerText?.replace(/\s+/g, ' ').trim(), "See more cat photos in our gallery.");
```

새 단락 안에 앵커 요소가 여전히 있어야 합니다.

```js
const P = document.querySelectorAll('p')[1];
assert.exists(P?.querySelector('a'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
