---
id: 587d78aa367417b2b2512aed
title: HTMLドキュメントのDoctypeを宣言する
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cra98AJ'
forumTopicId: 301095
dashedName: declare-the-doctype-of-an-html-document
---

# --description--

これまでのチャレンジでは、特定のHTML要素とその使い方を扱ってきました。しかし、ページ全体の構造を決めるいくつかの要素があり、それらはすべてのHTMLドキュメントに含める必要があります。

ドキュメントの最上部で、ブラウザーにページがどのバージョンのHTMLを使っているかを伝える必要があります。HTMLは進化し続けている言語で、定期的に更新されています。主要なブラウザーの多くは最新の仕様であるHTML5をサポートしていますが、古いウェブページは以前のバージョンを使っている場合があります。

この情報は、最初の行に`<!DOCTYPE ...>`タグを追加し、`...`の部分にHTMLのバージョンを指定することでブラウザーに伝えます。HTML5の場合は`<!DOCTYPE html>`を使います。

`!`と大文字の`DOCTYPE`は特に古いブラウザーにとって重要です。`html`は大文字・小文字を区別しません。

次に、残りのHTMLコードは`html`タグで囲む必要があります。開く`<html>`は`<!DOCTYPE html>`の直下に置き、閉じる`</html>`はページの最後に置きます。

以下はページ構造の例です。HTMLコードは2つの`html`タグの間に記述します。

```html
<!DOCTYPE html>
<html>

</html>
```

# --instructions--

コードエディターの空のHTMLドキュメントの最上部にHTML5用の`DOCTYPE`タグを追加してください。その下に、`html`の開閉タグを追加し、その中に`h1`要素を入れてください。見出しには任意のテキストを含めてかまいません。

# --hints--

コードには`<!DOCTYPE html>`タグを含めてください。

```js
assert.match(code,/<!DOCTYPE\s+?html\s*?>/gi);
```

`html`要素は1つだけにしてください。

```js
assert.lengthOf(document.querySelectorAll('html'), 1);
```

`html`タグは1つの`h1`要素を囲むようにしてください。

```js
assert.match(code,/<html>\s*?<h1>\s*?.*?\s*?<\/h1>\s*?<\/html>/gi);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <h1> Hello world </h1>
</html>
```
