---
id: bad87fee1348bd9aedf08816
title: アンカー要素で外部ページにリンクする
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/c8EkncB'
forumTopicId: 18226
dashedName: link-to-external-pages-with-anchor-elements
---

# --description--

`a`（*アンカー*）要素を使って、ウェブページの外部にあるコンテンツにリンクできます。

`a`要素には、`href`属性と呼ばれるリンク先のウェブアドレスが必要です。また、アンカーテキストも必要です。例を示します：

```html
<a href="https://www.freecodecamp.org">this links to freecodecamp.org</a>
```

するとブラウザーはテキスト`this links to freecodecamp.org`をクリックできるリンクとして表示します。そのリンクはウェブアドレス`https://www.freecodecamp.org`にリンクします。

# --instructions--

`a`要素を作成し、`https://www.freecatphotoapp.com`にリンクし、アンカーテキストを「cat photos」にしてください。

# --hints--

`a`要素のアンカーテキストは`cat photos`にしてください。

```js
assert.match(document.querySelector('a').textContent,/cat photos/gi);
```

`a`要素が`https://www.freecatphotoapp.com`にリンクする必要があります。

```js
assert.match(document.querySelector('a').getAttribute('href'),/^https?:\/\/(www\.)?freecatphotoapp\.com\/?$/i);
```

`a`要素には終了タグが必要です。

```js
assert.match(code,/<\/a>/g);
assert.strictEqual(code.match(/<\/a>/g).length,code.match(/<a/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>



  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  
  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
  
  <a href="https://www.freecatphotoapp.com">cat photos</a>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
