---
id: bad87fee1348bd9aedf08828
title: 番号付きリストを作成する
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cQ3B8TM'
forumTopicId: 16824
dashedName: create-an-ordered-list
---

# --description--

HTMLには、<dfn>番号付きリスト</dfn>を作成するための特別な要素があります。

番号付きリストは、開く`<ol>`要素から始まり、その後に任意の数の`<li>`要素が続きます。最後に、番号付きリストは`</ol>`タグで閉じられます。

例えば：

```html
<ol>
  <li>Garfield</li>
  <li>Sylvester</li>
</ol>
```

は`Garfield`と`Sylvester`の番号付きリストを作成します。

# --instructions--

猫が最も嫌うトップ3のものの番号付きリストを作成してください。

# --hints--

`Top 3 things cats hate:`の番号付きリストがあるはずです。

```js
const previousElement = document.querySelector('ol').previousElementSibling; 
assert.match(previousElement.textContent,/Top 3 things cats hate:/i);
```

`Things cats love:`の番号なしリストがあるはずです。

```js
const previousElement = document.querySelector('ul').previousElementSibling; 
assert.match(previousElement.textContent,/Things cats love:/i);
```

`ul`要素は1つだけにしてください。

```js
assert.lengthOf(document.querySelectorAll('ul'), 1);
```

`ol`要素は1つだけにしてください。

```js
assert.lengthOf(document.querySelectorAll('ol'), 1);
```

`li`要素の中に`ul`要素が3つあるはずです。

```js
assert.lengthOf(document.querySelectorAll('ul li'), 3);
```

`li`要素の中に`ol`要素が3つあるはずです。

```js
assert.lengthOf(document.querySelectorAll('ol li'), 3);
```

`ul`要素には終了タグが必要です。

```js
assert.match(code,/<\/ul>/g);
assert.strictEqual(code.match(/<\/ul>/g).length ,code.match(/<ul>/g).length);
```

`ol`要素には終了タグが必要です。

```js
assert.match(code,/<\/ol>/g);
assert.strictEqual(code.match(/<\/ol>/g).length ,code.match(/<ol>/g).length);
```

`li`要素には終了タグが必要です。

```js
assert.match(code,/<\/li>/g);
assert.match(code,/<li>/g);
assert.strictEqual(code.match(/<\/li>/g).length ,code.match(/<li>/g).length);
```

番号なしリストの`li`要素は空であってはいけません。

```js
[...document.querySelectorAll('ul li')].forEach((val) =>
  assert.isNotEmpty(__helpers.removeWhiteSpace(val.textContent))
);
```

番号付きリストの`li`要素は空であってはいけません。

```js
[...document.querySelectorAll('ol li')].forEach((val) =>
  assert.isNotEmpty(__helpers.removeWhiteSpace(val.textContent))
);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>

</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>hate 1</li>
    <li>hate 2</li>
    <li>hate 3</li>
  </ol>
</main>
```
