---
id: bad87fee1348bd9aedf08802
title: HTMLのコメント解除
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cBmG9T7'
forumTopicId: 18329
dashedName: uncomment-html
---

# --description--

コメントは、最終的にユーザーに表示される結果に影響を与えずに、他の開発者に向けてコード内にメモを残す方法です。

コメントはまた、コードを完全に削除せずに無効化する便利な方法でもあります。

HTMLのコメントは`<!--`で始まり、`-->`で終わります。

# --instructions--

`h1`、`h2`、`p`のコメントを解除してください。

# --hints--

`h1`の要素はコメントを解除するとページに表示されるはずです。

```js
assert.notEmpty(document.querySelectorAll('h1'));
```

`h2`の要素はコメントを解除するとページに表示されるはずです。

```js
assert.notEmpty(document.querySelectorAll('h2'));
```

`p`の要素はコメントを解除するとページに表示されるはずです。

```js
assert.notEmpty(document.querySelectorAll('p'));
```

ページにコメントタグの残り（つまり`-->`）が表示されないようにしてください。

```js
const elements = document.querySelectorAll('*');
const potentialComments = Array.from(elements).filter(el => el.textContent.includes('-->'));
assert.notExists(potentialComments[1])
```

# --seed--

## --seed-contents--

```html
<!--
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
-->
```

# --solutions--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```
