---
id: bad87fee1348bd9aedf08834
title: ラジオボタンのセットを作成する
challengeType: 0
forumTopicId: 16822
dashedName: create-a-set-of-radio-buttons
---

# --description--

複数の選択肢の中からユーザーに1つだけ答えてほしい質問には、<dfn>ラジオボタン</dfn>を使うことができます。

ラジオボタンは`input`の一種です。

それぞれのラジオボタンは、自分専用の`label`要素の中にネストすることができます。`input`要素の中に`label`要素を入れることで、そのラジオボタンの入力と周囲のラベル要素が自動的に関連付けられます。

関連するすべてのラジオボタンは、同じ`name`属性を持つ必要があり、これによってラジオボタングループが作られます。ラジオグループを作ることで、どれか1つのラジオボタンを選択すると、同じグループ内の他のボタンは自動的に選択解除され、ユーザーが1つだけ答えを選べるようになります。

ラジオボタンの例を示します：

```html
<label> 
  <input type="radio" name="indoor-outdoor">Indoor 
</label>
```

`for`要素に`label`属性を設定し、その値を`id`要素の`input`属性の値と一致させることがベストプラクティスとされています。これにより支援技術がラベルと関連する`input`要素の間にリンク関係を作れます。例えば：

```html
<input id="indoor" type="radio" name="indoor-outdoor">
<label for="indoor">Indoor</label>
```

また、`input`要素を`label`タグの中にネストすることもできます：

```html
<label for="indoor"> 
  <input id="indoor" type="radio" name="indoor-outdoor">Indoor 
</label>
```

# --instructions--

フォームにラジオボタンを2つ追加してください。それぞれは自分専用の`label`要素の中にネストしてください。1つは`indoor`の選択肢を持ち、もう1つは`outdoor`の選択肢を持つようにします。両方とも`name`属性に`indoor-outdoor`を設定してラジオグループを作ってください。

# --hints--

ページには2つの`radio`ボタン要素が必要です。

```js
assert.lengthOf(document.querySelectorAll('input[type="radio"]'),2);
```

ラジオボタンには`name`属性として`indoor-outdoor`を設定してください。

```js
assert.lengthOf([...document.querySelectorAll('input[type="radio"]')].filter(input => input.name === "indoor-outdoor"),2);
```

2つのラジオボタン要素はそれぞれ自分専用の`label`要素の中にネストしてください。

```js
assert.lengthOf(document.querySelectorAll('label > input[type="radio"]:only-child'),2);
```

`label`要素には終了タグが必要です。

```js
assert.match(code,/<\/label>/g);
assert.match(code,/<label/g);
assert.strictEqual(code.match(/<\/label>/g).length,code.match(/<label/g).length);
```

ラジオボタンのうち1つにはラベル`indoor`を付けてください。

```js
const labelTexts = [...document.querySelectorAll('label')].map(label => label.textContent);
const hasMatch = labelTexts.some(label =>  label.match(/indoor/gi));
assert.isTrue(hasMatch); 
```

もう1つのラジオボタンにはラベル`outdoor`を付けてください。

```js
const labelTexts = [...document.querySelectorAll('label')].map(label => label.textContent);
const hasMatch = labelTexts.some(label => label.match(/outdoor/gi));
assert.isTrue(hasMatch); 
```

ラジオボタン要素はすべて`form`タグの中に追加してください。

```js
const labelParent = document.querySelector('label').parentNode;
assert.strictEqual(labelParent.tagName,'FORM');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
   <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
