---
id: bad87fee1348bd9aedf08812
title: ウェブサイトに画像を追加する
challengeType: 0
forumTopicId: 16640
dashedName: add-images-to-your-website
---

# --description--

ウェブサイトに画像を追加するには、`img` 要素を使い、`src` 属性で特定の画像のURLを指定します。

例は以下の通りです：

```html
<img src="https://www.freecatphotoapp.com/your-image.jpg">
```

`img` は空要素であることに注意してください。

すべての `img` 要素には `alt` 属性が**必要**です。`alt` 属性内のテキストはスクリーンリーダーのアクセシビリティ向上に使われ、画像が読み込めなかった場合に表示されます。

**注意：** 画像が純粋に装飾目的の場合は、空の `alt` 属性を使うのがベストプラクティスです。

理想的には `alt` 属性には必要な場合を除き特殊文字を含めないようにしてください。

上記の `alt` の例に `img` 属性を追加してみましょう：

```html
<img src="https://www.freecatphotoapp.com/your-image.jpg" alt="freeCodeCamp logo">
```

# --instructions--

ウェブサイトに画像を追加してみましょう：

既存の `main` 要素内で、既存の `img` 要素の前に `p` 要素を挿入してください。

次に、`src` 属性を設定して、URL `https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg` を指すようにしてください。

最後に、`img` 要素に適切なテキストを持つ `alt` 属性を忘れずに付けてください。

# --hints--

ページに画像要素があることを確認してください。

```js
assert.exists(document.querySelector('img'));
```

画像には子猫の画像を指す `src` 属性が必要です。

```js
const url = document.querySelector('img').getAttribute('src');
assert.match(url,/^https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg$/i);
```

画像要素の `alt` 属性は空であってはいけません。

```js
assert.exists(document.querySelector('img').getAttribute('alt'));
assert.isNotEmpty(document.querySelector('img').getAttribute('alt'));
assert.match(__helpers.removeWhiteSpace(code),/<(?:img|IMG)\S*alt=(['"])(?!\1|>)\S+\1\S*\/?>/)
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>


  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
