---
id: bad87fee1348bd9aedf08801
title: 段落要素で情報を伝える
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/ceZ7DtN'
forumTopicId: 18202
dashedName: inform-with-the-paragraph-element
---

# --description--

`p` 要素はウェブサイトの段落テキストに推奨される要素です。`p` は「段落」の略です。

段落要素は次のように作成できます。

```html
<p>I'm a p tag!</p>
```

# --instructions--

`p` 要素を `h2` 要素の下に作成し、テキストに `Hello Paragraph` を入れてください。

**注意:** 慣例として、すべてのHTMLタグは小文字で書きます。例えば `<p></p>` であり、`<P></P>` ではありません。

# --hints--

コードには有効な `p` 要素が含まれている必要があります。

```js
assert.lengthOf(document.querySelectorAll('p'),1);
```

`p` 要素にはテキスト `Hello Paragraph` が含まれている必要があります。

```js
assert.match(document.querySelector('p').textContent,/hello(\s)+paragraph/gi);
```

`p` 要素には終了タグが必要です。

```js
assert.match(code,/<\/p>/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>
<h2>CatPhotoApp</h2>
```

# --solutions--

```html
<h1>Hello World</h1>
<h2>CatPhotoApp</h2>
<p>Hello Paragraph</p>
```
