---
id: bad87fee1348bd9aedf08827
title: 番号なしリストを作成する
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cDKVPuv'
forumTopicId: 16814
dashedName: create-a-bulleted-unordered-list
---

# --description--

HTMLには、<dfn>番号なしリスト</dfn>、つまり箇条書きスタイルのリストを作成するための特別な要素があります。

番号なしリストは、開く`<ul>`要素から始まり、その後に任意の数の`<li>`要素が続きます。最後に、番号なしリストは`</ul>`で閉じられます。

例えば：

```html
<ul>
  <li>milk</li>
  <li>cheese</li>
</ul>
```

は`milk`と`cheese`の箇条書きスタイルのリストを作成します。

# --instructions--

最後の2つの`p`要素を削除し、ページの下部に猫が好きな3つのものの番号なしリストを作成してください。

# --hints--

`ul`要素を作成してください。

```js
assert.isNotEmpty(document.querySelectorAll('ul'));
```

`li`要素の中に`ul`要素が3つあるはずです。

```js
assert.lengthOf(document.querySelectorAll('ul li'),3);
```

`ul`要素には終了タグが必要です。

```js
assert.match(code,/<\/ul>/gi);
assert.match(code,/<ul/gi);
assert.strictEqual(code.match(/<\/ul>/gi).length,code.match(/<ul/gi).length);
```

`li`要素には終了タグが必要です。

```js
assert.match(code,/<\/li>/gi);
assert.match(code,/<li[\s>]/gi);
assert.strictEqual(code.match(/<\/li>/gi).length,code.match(/<li[\s>]/gi).length);
```

`li`要素には空の文字列や空白だけを含めないでください。

```js
assert.isEmpty([...document.querySelectorAll('ul li')].filter((item) => item.textContent.trim() === ""));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <ul>
    <li>milk</li>
    <li>mice</li>
    <li>catnip</li>
  </ul>
</main>
```
