---
id: bad87fee1348bd9aede08817
title: 段落内にアンカー要素をネストする
challengeType: 0
forumTopicId: 18244
dashedName: nest-an-anchor-element-within-a-paragraph
---

# --description--

リンクは他のテキスト要素の中にネストすることができます。

```html
<p>
  Here's a <a target="_blank" href="https://www.freecodecamp.org"> link to www.freecodecamp.org</a> for you to follow.
</p>
```

例を分解してみましょう。通常のテキストは`p`要素で囲まれています：

```html
<p> Here's a ... for you to follow. </p>
```

次に*アンカー*要素`<a>`（終了タグ`</a>`が必要です）があります：  

```html
<a> ... </a>
```

`target`はリンクをどこで開くかを指定するアンカータグの属性です。値`_blank`はリンクを新しいタブで開くことを指定します。`href`はリンクのURLアドレスを含むアンカータグの属性です：  

```html
<a href="https://www.freecodecamp.org" target="_blank"> ... </a>
```

`link to www.freecodecamp.org`要素内のテキスト`a`は<dfn>アンカーテキスト</dfn>と呼ばれ、クリックできるリンクとして表示されます：

```html
<a href=" ... " target="...">link to freecodecamp.org</a>
```

例の最終的な出力は次のようになります：  

こちらは<a href="https://www.freecodecamp.org" target="_blank">www.freecodecamp.orgへのリンク</a>です。

# --instructions--

既存の`a`要素を新しい`p`要素の中にネストしてください。新しいアンカータグは作成しないでください。新しい段落は`View more cat photos`というテキストを含み、その中の`cat photos`がリンクで、残りはプレーンテキストである必要があります。

# --hints--

`a`要素は1つだけにしてください。

```js
assert.lengthOf(document.querySelectorAll('a'), 1 );
```

`a`要素は「`https://www.freecatphotoapp.com`」にリンクする必要があります。

```js
assert.lengthOf(document.querySelectorAll('a[href="https://www.freecatphotoapp.com"]'),1);
```

`a`要素のアンカーテキストは`cat photos`であるべきです。

```js
assert.match(document.querySelector('a').textContent,/cat\sphotos/gi)
```

新しい`p`要素を作成してください。HTMLコード内には合計で少なくとも3つの`p`タグが必要です。

```js
assert.lengthOf(document.querySelectorAll('p'),3)
```

`a`要素は新しい`p`要素の中にネストされている必要があります。

```js
const anchorParent = document.querySelector('a[href="https://www.freecatphotoapp.com"]').parentNode;
assert.strictEqual(anchorParent.tagName,"P")
```

`p`要素は`View more `というテキスト（後にスペースあり）を含むべきです。

```js
const textContent = document.querySelector('a[href="https://www.freecatphotoapp.com"]').parentNode.textContent;
assert.match(textContent,/View\smore\s/gi);
```

`a`要素は`View more`というテキストを<em>含んではいけません</em>。

```js
assert.notMatch(document.querySelector('a').textContent,/View\smore/gi); 
```

`p`要素にはそれぞれ終了タグが必要です。

```js
assert.match(code,/<\/p>/g);
assert.match(code,/<p/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

`a`要素にはそれぞれ終了タグが必要です。

```js
assert.match(code,/<\/a>/g);
assert.match(code,/<a/g);
assert.strictEqual(code.match(/<\/a>/g).length,code.match(/<a/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>View more <a target="_blank" href="https://www.freecatphotoapp.com">cat photos</a></p>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
