---
id: bad87fee1348bd9aedf08804
title: HTMLをコメントアウトする
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cGyGbca'
forumTopicId: 16782
dashedName: comment-out-html
---

# --description--

コメントを開始するには`<!--`を使い、コメントを終了するには`-->`を使う必要があることを覚えておいてください。

ここでは、`h2`要素が始まる前にコメントを終了する必要があります。

# --instructions--

`h1`要素と`p`要素をコメントアウトしてください。ただし、`h2`要素はコメントアウトしないでください。

# --hints--

`h1`要素はページに表示されないようにコメントアウトしてください。

```js
assert.isEmpty(document.querySelectorAll('h1'));
```

`h2`要素はページに表示されるようにコメントアウトしないでください。

```js
assert.isNotEmpty(document.querySelectorAll('h2'));
```

`p`要素はページに表示されないようにコメントアウトしてください。

```js
assert.isEmpty(document.querySelectorAll('p'));
```

コメントはすべて`-->`で閉じる必要があります。

```js
assert.isAbove(code.match(/[^fc]-->/g).length, 1);
```

コード内の`h1`、`h2`、`p`の順序は変更しないでください。

```js
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[0],'<h1>');
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[1],'<h2>');
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[2],'<p>');
```

# --seed--

## --seed-contents--

```html
<!--
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
-->
```

# --solutions--

```html
<!--<h1>Hello World</h1>-->
<h2>CatPhotoApp</h2> 
<!--<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p> -->
```
