---
id: bad87fee1348bd9aedf08833
title: プレースホルダーテキストで空欄を埋める
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cgR7Dc7'
forumTopicId: 18178
dashedName: fill-in-the-blank-with-placeholder-text
---

# --description--

ウェブ開発者は伝統的に<dfn>lorem ipsum テキスト</dfn>をプレースホルダーテキストとして使っています。lorem ipsum テキストは古代ローマのキケロの有名な文章からランダムに抜き出されたものです。

lorem ipsum テキストは16世紀から組版業者によってプレースホルダーテキストとして使われており、この伝統はウェブでも続いています。

さて、5世紀も十分長いです。CatPhotoAppを作っているので、「kitty ipsum」テキストを使いましょう。

# --instructions--

`p` 要素の中のテキストを、この kitty ipsum テキストの最初の数語に置き換えてください：`Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.`

# --hints--

`p` 要素には、提供された「kitty ipsum」テキストの最初の数語を入れてください。

```js
assert.match(document.querySelector('p').textContent,/Kitty(\s)+ipsum/gi);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Hello Paragraph</p>
```

# --solutions--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff</p>
```
