---
id: 686e64be3cd9db25282a956d
title: تصميم معرض تصوير فوتوغرافي
challengeType: 25
dashedName: lab-photography-exhibit
demoType: onClick
saveSubmissionToDB: true
---

# --description--

في هذا المختبر، ستصمم تخطيط معرض تصوير فوتوغرافي باستخدام فئات الأدوات المساعدة من Tailwind CSS.

يمكنك استخدام عناوين URL التالية للصور إذا أردت:

- https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg
- https://cdn.freecodecamp.org/curriculum/labs/alps.jpg
- https://cdn.freecodecamp.org/curriculum/labs/sea.jpg

**الهدف:** حقق قصص المستخدم أدناه واجتز جميع الاختبارات لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك عنصر بفئة اسمها `main-container`.
2. يجب أن يحتوي عنصر `.main-container` أيضًا على فئات الأدوات المساعدة التالية من Tailwind CSS:
   - فئة الأدوات المساعدة الصحيحة لإنشاء حاوية شبكة CSS.
   - فئة أدوات مساعدة لتحديد عدد الأعمدة داخل الشبكة. يجب استخدام فئة الأدوات المساعدة التي تستخدم عددًا ثابتًا من الأعمدة. مثال: `grid-cols-<number>`.
   - فئة أدوات مساعدة لتعيين الفجوة بين الصفوف والأعمدة لعناصر الشبكة.
3. داخل عنصر `.main-container` يجب أن يكون هناك ثلاثة عناصر على الأقل، كل منها بفئة اسمها `card`.
4. يجب أن تحتوي كل عناصر `.card` على فئات الأدوات المساعدة التالية من Tailwind CSS:
   - فئة أدوات مساعدة لتعيين حدود مستديرة من اختيارك.
   - فئة أدوات مساعدة لتعيين الحشو من اختيارك. يجب استخدام هذا التنسيق `p-<number>`.
5. داخل كل عنصر من عناصر `.card` يجب أن يحتوي على العناصر التالية:
   - عنصر صورة مع خاصيتي `src` و `alt`.
   - عنصر بفئة اسمها `subheading`.
   - عنصر بفئة اسمها `description`.
6. يجب أن تحتوي كل عناصر الصورة على فئة أدوات مساعدة لتعيين حدود مستديرة من اختيارك.
7. يجب أن تستخدم كل عناصر `.subheading` فئات أدوات مساعدة لتعيين وزن الخط وحجمه حسب اختيارك.
8. يجب أن تستخدم كل عناصر `.description` فئة أدوات مساعدة لتعيين حجم الخط حسب اختيارك.

# --hints--

يجب أن يكون لديك عنصر بفئة اسمها `main-container`.

```js
assert.exists(document.querySelector(".main-container"));
```

يجب أن يستخدم عنصر `.main-container` فئة الأدوات المساعدة الصحيحة من Tailwind CSS لإنشاء حاوية شبكة CSS.

```js
const mainContainer = document.querySelector(".main-container");
assert.isTrue(mainContainer.classList.contains("grid"));
```

يجب أن يستخدم عنصر `.main-container` فئة أدوات مساعدة لتحديد عدد الأعمدة داخل الشبكة.

```js
const mainContainer = document.querySelector(".main-container");
// Grid classes can vary, so this check is for the presence of any grid-cols class
const hasGridCols = [...mainContainer.classList].some(className => className.startsWith("grid-cols-"));
assert.isTrue(hasGridCols);
```

يجب أن يستخدم عنصر `.main-container` فئة أدوات مساعدة لتعيين الفجوة بين الصفوف والأعمدة لعناصر الشبكة.

```js
const mainContainer = document.querySelector(".main-container");
// Check for any gap class including gap-x and gap-y

const hasGap = [...mainContainer.classList].some(c => c.startsWith("gap-")) ||
  (classList.some(c => c.startsWith("gap-x-")) &&
   classList.some(c => c.startsWith("gap-y-")));

assert.isTrue(hasGap);
```

يجب أن يحتوي عنصر `.main-container` على ثلاثة عناصر على الأقل، كل منها بفئة اسمها `card`.

```js
const cards = document.querySelectorAll(".main-container .card");
assert.isAtLeast(cards.length, 3);
```

يجب أن تحتوي كل عناصر `.card` على فئة أدوات مساعدة لتعيين حدود مستديرة.

```js
const cards = document.querySelectorAll(".main-container .card");
assert.isNotEmpty(cards);

const predefinedRoundedCorners = [
  "rounded-none",
  "rounded-sm",
  "rounded",
  "rounded-md",
  "rounded-lg",
  "rounded-xl",
  "rounded-2xl",
  "rounded-3xl",
  "rounded-full"
];

cards.forEach(card => {
  const hasRoundedCorners = predefinedRoundedCorners.some(corner => card.classList.contains(corner));
  assert.isTrue(hasRoundedCorners);
})
```

يجب أن تحتوي كل عناصر `.card` على فئة أدوات مساعدة لتعيين الحشو.

```js
const cards = document.querySelectorAll(".main-container .card");
assert.isNotEmpty(cards);
const hasPadding = [...cards].every(card => {
  return [...card.classList].some(c => c.startsWith("p-"));
});
assert.isTrue(hasPadding);
```

يجب أن يحتوي كل عنصر من عناصر `.card` على عنصر صورة.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isAtLeast(images.length, 3);
```

يجب أن تحتوي كل عناصر `img` على خاصية `src`.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isAtLeast(images.length, 3);
images.forEach(img => assert.isTrue(img.hasAttribute("src")));
images.forEach(img => assert.isTrue(img.getAttribute("src").length > 0));
```

يجب أن تحتوي كل عناصر `img` على خاصية `alt`.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isAtLeast(images.length, 3);
images.forEach(img => assert.isTrue(img.hasAttribute("alt")));
images.forEach(img => assert.isTrue(img.getAttribute("alt").length > 0));
```

يجب أن يحتوي كل عنصر من عناصر `.card` على عنصر بفئة اسمها `subheading`.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
assert.isAtLeast(subheadings.length, 3);
```

يجب أن يحتوي كل عنصر من عناصر `.subheading` على محتوى نصي.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
// Ensure we are not dealing with an empty NodeList
assert.isNotEmpty(subheadings);

subheadings.forEach(subheading => {
  assert.isNotEmpty(subheading.textContent.trim());
});
```

يجب أن يحتوي كل عنصر من عناصر `.card` على عنصر بفئة اسمها `description`.

```js
const descriptions = document.querySelectorAll(".main-container .card .description");
assert.isAtLeast(descriptions.length, 3);
```

يجب أن يحتوي كل عنصر من عناصر `.description` على محتوى نصي.

```js
const descriptions = document.querySelectorAll(".main-container .card .description");
// Ensure we are not dealing with an empty NodeList
assert.isNotEmpty(descriptions);

descriptions.forEach(description => {
  assert.isNotEmpty(description.textContent.trim());
});
```

يجب أن تحتوي كل عناصر `img` على فئة أدوات مساعدة لتعيين حدود مستديرة.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isNotEmpty(images);
const predefinedRoundedCorners = [
  "rounded-none",
  "rounded-sm",
  "rounded",
  "rounded-md",
  "rounded-lg",
  "rounded-xl",
  "rounded-2xl",
  "rounded-3xl",
  "rounded-full"
];

images.forEach(img => {
  const hasRoundedCorners = predefinedRoundedCorners.some(corner => img.classList.contains(corner));
  assert.isTrue(hasRoundedCorners);
});
```

يجب أن تستخدم كل عناصر `.subheading` فئات أدوات مساعدة لتعيين وزن الخط.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
assert.isNotEmpty(subheadings);
const predefinedFontWeights = [
  "font-thin",
  "font-extralight",
  "font-light",
  "font-normal",
  "font-medium",
  "font-semibold",
  "font-bold",
  "font-extrabold",
  "font-black"
];

subheadings.forEach(subheading => {
  const hasFontWeight = predefinedFontWeights.some(weight => subheading.classList.contains(weight));
  assert.isTrue(hasFontWeight);
});
```

يجب أن تستخدم كل عناصر `.subheading` فئات أدوات مساعدة لتعيين حجم الخط.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
assert.isNotEmpty(subheadings);
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

subheadings.forEach(subheading => {
  const hasFontSize = predefinedSizes.some(size => subheading.classList.contains(size));
  assert.isTrue(hasFontSize);
});
```

يجب أن تستخدم كل عناصر `.description` فئة أدوات مساعدة لتعيين حجم الخط.

```js
const descriptions = document.querySelectorAll(".description");
assert.isNotEmpty(descriptions);
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

descriptions.forEach(description => {
  const hasFontSize = predefinedSizes.some(size => description.classList.contains(size));
  assert.isTrue(hasFontSize);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Photography Exhibit</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
 
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Photography Exhibit</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-purple-800 text-yellow-400 min-h-screen">
    <h1 class="text-center text-4xl font-bold my-12">Photography Exhibit</h1>

    <main class="main-container grid grid-cols-1 md:grid-cols-3 gap-8 px-4">
      <figure
        class="card bg-purple-900 border border-yellow-400 rounded-lg p-4"
      >
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="Ancient Colosseum"
          class="h-48 w-full object-cover rounded-md mb-4"
        />
        <figcaption>
          <h2 class="subheading text-lg font-semibold mb-2">Colosseum, Rome</h2>
          <p class="description text-sm mb-4">
            There is so much history and beauty. The light hits it just right.
          </p>
          <div class="grid grid-cols-2 gap-4">
            <button
              class="bg-purple-800 hover:bg-purple-700 border border-yellow-400 py-1 px-2 rounded-md"
            >
              Like
            </button>
            <a
              href="#"
              class="bg-yellow-400 hover:bg-yellow-300 text-purple-900 font-semibold py-1 px-2 rounded-md text-center"
            >
              View Details
            </a>
          </div>
        </figcaption>
      </figure>

      <figure
        class="card bg-purple-900 border border-yellow-400 rounded-lg p-4"
      >
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="Alpine Mountains"
          class="h-48 w-full object-cover rounded-md mb-4"
        />
        <figcaption>
          <h2 class="subheading text-lg font-semibold mb-2">The Alps</h2>
          <p class="description text-sm mb-4">
            The mountains go on forever and the view is breathtaking.
          </p>
          <div class="grid grid-cols-2 gap-4">
            <button
              class="bg-purple-800 hover:bg-purple-700 border border-yellow-400 py-1 px-2 rounded-md"
            >
              Like
            </button>
            <a
              href="#"
              class="bg-yellow-400 hover:bg-yellow-300 text-purple-900 font-semibold py-1 px-2 rounded-md text-center"
            >
              View Details
            </a>
          </div>
        </figcaption>
      </figure>

      <figure
        class="card bg-purple-900 border border-yellow-400 rounded-lg p-4"
      >
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="Calm Sea"
          class="h-48 w-full object-cover rounded-md mb-4"
        />
        <figcaption>
          <h2 class="subheading text-lg font-semibold mb-2">Seascape</h2>
          <p class="description text-sm mb-4">
            The water was still, the air was cool and the feeling is pure peace.
          </p>
          <div class="grid grid-cols-2 gap-4">
            <button
              class="bg-purple-800 hover:bg-purple-700 border border-yellow-400 py-1 px-2 rounded-md"
            >
              Like
            </button>
            <a
              href="#"
              class="bg-yellow-400 hover:bg-yellow-300 text-purple-900 font-semibold py-1 px-2 rounded-md text-center"
            >
              View Details
            </a>
          </div>
        </figcaption>
      </figure>
    </main>
  </body>
</html>
```
