---
id: 60f02e7361b68405e27b62a5
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在這個工作坊中，你將學習如何透過設計註冊表單來為表單套用樣式。

所有的 HTML 標準範本（`DOCTYPE`、`html`、`head` 和 `body`）都已為你提供。

在 `body` 中，使用帶有文字 `Registration Form` 的 `h1` 元素為表單提供背景關係。

# --hints--

你應該在 `body` 中添加 `h1`。

```js
assert.exists(document.querySelector('body > h1'));
```

你應該給 `h1` 設定文字 `Registration Form`。

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
