---
id: 60f030d388cb74067cf291c3
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在標題下方，使用以下文字於段落元素中以鼓勵使用者註冊：

```md
Please fill out this form with the required information
```

# --hints--

你應該在 `body` 中添加一個 `p` 元素。

```js
assert.exists(document.querySelector('body > p'));
```

你應該在 `h1` 下方添加 `p` 元素。

```js
assert.exists(document.querySelector('h1 + p'));
```

你應該給 `p` 元素一個 `Please fill out this form with the required information` 的文字。

```js
assert.equal(document.querySelector('p')?.innerText, 'Please fill out this form with the required information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
--fcc-editable-region--
<body>
  <h1>Registration Form</h1>
  
</body>
--fcc-editable-region--
</html>
```
