---
id: 60f1a9cbd23023082e149fee
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

那樣比較好。現在，將 `body` 的 `background-color` 改為 `#1b1b32`，讓背景更舒適。然後，為了看清楚文字，將 `color` 改為 `#f5f6f7`。

# --hints--

你應該在 `body` 元素選擇器中添加 `background-color`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

你應該為 `background-color` 指定一個值 `#1b1b32`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

你應該在 `body` 元素選擇器內添加 `color`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

你應該給 `color` 一個 `#f5f6f7` 的值。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
