---
id: 60f1922fcbd2410527b3bd89
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

在之前的課程中，你學會了如何使用 viewport 單位。

請記住，`vh` 單位代表視窗高度，且相當於視窗 `height` 的 1%。這使它相對於視窗高度。

是時候用一些 CSS 美化專案了。首先，給 `body` 設定 `width` 為 `100%`，以及 `height` 為 `100vh`。

# --hints--

你應該使用 `body` 元素選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

你應該給 `body` 一個 `width` 為 `100%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

你應該給 `body` 一個 `height` 為 `100vh`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
