---
id: 60f5c3e399ff1a05629964e4
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

如標題所示，你正在建立一個表單。所以，在 `p` 元素之後，插入一個具有 `action` 屬性且標的為 `https://register-demo.freecodecamp.org` 的 `form`。

# --hints--

你應該在 `p` 元素旁邊添加一個 `form` 元素。

```js
assert.exists(document.querySelector('p + form'));
```

你應該為 `form` 指定一個 `action` 屬性。

```js
// Default action points to window location
assert.notEqual(document.querySelector('form')?.action, window?.location?.href);
```

你應該給 `action` 一個值為 `https://register-demo.freecodecamp.org`。

```js
assert.equal(document.querySelector('form')?.action, 'https://register-demo.freecodecamp.org/');
```

你的 `form` 元素應該有一個結束標籤 `</form>`。

```js
assert.match(code, /<\/form\>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
