---
id: 60f5cb8875ab6a0610f05071
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

由於表單將有三個不同的區段，請在 `form` 元素內添加三個 `fieldset` 元素。

# --hints--

你應該添加三個 `fieldset` 元素。

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

三個 `fieldset` 元素應該位於 `form` 內。

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

這三個 `fieldset` 元素應該是同層元素。

```js
assert.exists(document.querySelector('fieldset + fieldset + fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
    
--fcc-editable-region--
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
