---
id: 60f5c3e399ff1a05629964e4
title: 步驟6
challengeType: 0
dashedName: step-6
---

# --description--

正如標題所述，你正在創建表單。 因此，在 `p` 元素之後，插入一個 `form` 元素，其 `action` 屬性值爲 `https://register-demo.freecodecamp.org`。

# --hints--

你應該在 `p` 元素之後添加一個 `form` 元素。

```js
assert.exists(document.querySelector('p + form'));
```

你應該給 `form` 元素添加一個 `action` 屬性。

```js
// Default action points to window location
assert.notEqual(document.querySelector('form')?.action, window?.location?.href);
```

你應該將 `action` 的值設爲 `https://register-demo.freecodecamp.org`。

```js
assert.equal(document.querySelector('form')?.action, 'https://register-demo.freecodecamp.org/');
```

你的 `form` 元素應該沒有結束標籤 `</form>`。

```js
assert.match(code, /<\/form\>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
--fcc-editable-region--

--fcc-editable-region--
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
