---
id: 60f81616cff80508badf9ad5
title: 第 16 步
challengeType: 0
dashedName: step-16
---

# --description--

此時，你應該能夠提交表單。然而，你可能會注意到沒有太大變化。

為了讓表單更具互動性，請將 `required` 屬性添加到第一個 `fieldset` 中的 `input` 元素。

現在，如果你嘗試在未填寫必填欄位的情況下提交表單，你將會看到錯誤訊息。

# --hints--

你應該為第一個 `input` 元素添加 `required` 屬性。

```js
assert.isTrue(document.querySelector('input')?.required);
```

你應該為第二個 `input` 元素添加 `required` 屬性。

```js
assert.isTrue(document.querySelectorAll('input')?.[1]?.required);
```

你應該為第三個 `input` 元素添加 `required` 屬性。

```js
assert.isTrue(document.querySelectorAll('input')?.[2]?.required);
```

你應該為第四個 `input` 元素添加 `required` 屬性。

```js
assert.isTrue(document.querySelectorAll('input')?.[3]?.required);
```

你不應該給型別為 `submit` 的 `input` 添加 `required` 屬性。

```js
assert.isFalse(document.querySelector('input[type="submit"]')?.required);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
