---
id: 60f803d5241e6a0433a523a1
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

在之前的課程中，你學會了如何使用 `rem` 單位。請記住，`rem` 單位代表根 `em`，並且是相對於 `html` 元素的字型大小。

由於 `label` 元素預設為行內，因此它們都會並排顯示在同一行，使其文字難以閱讀。

為了讓它們出現在不同的行上，為 `label` 元素添加 `display: block`，並添加 `0.5rem 0` 的 `margin`，以便彼此分隔。

# --hints--

你應該使用 `label` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('label'));
```

你應該為你的 `display` 屬性添加值為 `block`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.display, 'block');
```

你應該添加值為 `0.5rem 0` 的 `margin` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.margin, '0.5rem 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label>Enter Your First Name:</label>
      <label>Enter Your Last Name:</label>
      <label>Enter Your Email:</label>
      <label>Create a New Password:</label>
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--
```
