---
id: 60f84ec41116b209c280ba91
title: 步驟 18
challengeType: 0
dashedName: step-18
---

# --description--

使用 `type="password"` 時，你可以使用 `pattern` 屬性來定義密碼必須符合的正規表達式，才能被視為有效。

為你的密碼 `input` 元素添加 `pattern` 屬性，以要求輸入符合：`[a-z0-5]{8,}`

上述是一個正規表達式，匹配八個或以上的小寫字母或數字 `0` 到 `5`。然後，移除 `minlength` 屬性，並試用看看。

# --hints--

你應該為密碼 `input` 元素添加 `pattern` 屬性。

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.pattern);
```

你應該為 `pattern` 屬性指定一個值 `[a-z0-5]{8,}`。

```js
assert.equal(document.querySelector('input[type="password"]')?.pattern, '[a-z0-5]{8,}');
```

你應該從密碼 `input` 元素中移除 `minlength` 屬性。

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, -1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
--fcc-editable-region--
      <label for="new-password">Create a New Password: <input id="new-password" type="password" minlength="8" required /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
