---
id: 60f805f813eaf2049bc2ceea
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

將每個 `input` 元素巢狀於每個 `label` 中。請確保在 `label` 文字之後添加每個 `input`，並在冒號後加一個空格。

# --hints--

你應該為你的 `fieldset` 元素添加四個 `input` 元素。

```js
assert.equal(document.querySelectorAll('fieldset input')?.length, 4);
```

你應該將 `input` 元素巢狀於 `label` 元素內。

```js
assert.equal(document.querySelectorAll('label input')?.length, 4);
```

你應該在 `label` 文字 `Enter Your First Name:` 後添加第一個 `input`，並在冒號後包含一個空格。

```js
const query = /^Enter Your First Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[0]?.innerHTML.trim(), query);
```

你應該在 `label` 文字 `Enter Your Last Name:` 之後添加第二個 `input`，並在冒號後包含一個空格。

```js
const query = /^Enter Your Last Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[1]?.innerHTML.trim(), query);
```

你應該在 `label` 文字 `Enter Your Email:` 之後添加第三個 `input`，並在冒號後包含一個空格。

```js
const query = /^Enter Your Email:\s+<input/
assert.match(document.querySelectorAll('label')?.[2]?.innerHTML.trim(), query);
```

你應該在 `label` 文字 `Create a New Password:` 後面添加第四個 `input`，並在冒號後面加一個空格。

```js
const query = /^Create a New Password:\s+<input/
assert.match(document.querySelectorAll('label')?.[3]?.innerHTML.trim(), query);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label>Enter Your First Name:</label>
      <label>Enter Your Last Name:</label>
      <label>Enter Your Email:</label>
      <label>Create a New Password:</label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
