---
id: 60f80e0081e0f2052ae5b505
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

指定 `input` 元素的 `type` 屬性對瀏覽器了解應該預期何種資料非常重要。如果未指定 `type`，瀏覽器將預設為 `text`。

為前兩個 `input` 元素添加 `type` 屬性為 `text`，第三個添加 `type` 屬性為 `email`，第四個添加 `type` 屬性為 `password`。

`email` 型別只允許在網域中包含 `@` 和 `.` 的電子郵件。  
`password` 型別會隱藏輸入內容，並在網站未使用 HTTPS 時發出警告。

# --hints--

你應該為第一個 `input` 元素給予一個 `type` 屬性，值為 `text`。

```js
assert.equal(document.querySelector('input')?.type, 'text');
```

你應該給第二個 `input` 元素一個 `type` 屬性，值為 `text`。

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'text');
```

你應該為第三個 `input` 元素添加 `type` 屬性，並設定為 `email`。

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'email');
```

你應該為第四個 `input` 元素給予 `type` 屬性為 `password`。

```js
assert.equal(document.querySelectorAll('input')?.[3]?.type, 'password');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label for="first-name">Enter Your First Name: <input id="first-name" /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" /></label>
      <label for="email">Enter Your Email: <input id="email" /></label>
      <label for="new-password">Create a New Password: <input id="new-password" /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
