---
id: 60f81167d0d4910809f88945
title: 步驟15
challengeType: 0
dashedName: step-15
---

# --description--

`type` 爲 `submit` 的第一個 `input` 元素會自動設置爲提交其最近的 `form` 父元素。

爲了處理表格提交，在最後一個 `fieldset` 元素之後添加一個 `input` 元素，`type` 屬性設置爲 `submit`，`value` 屬性設置爲 `Submit`。

# --hints--

你應該在最後一個 `fieldset` 元素後添加 `input` 元素。

```js
assert.exists(document.querySelectorAll('fieldset')?.[2]?.nextElementSibling?.tagName, 'input');
```

你應該給 `input` 元素一個 `type` 屬性，值爲 `submit`。

```js
assert.exists(document.querySelector('fieldset + input[type="submit"]'));
```

你應該給 `input` 元素一個 `value` 屬性，值爲 `Submit`。

```js
assert.exists(document.querySelector('fieldset + input[value="Submit"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>

--fcc-editable-region--
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```

