---
id: 60facde2d0dc61085b41063f
title: 步驟 34
challengeType: 0
dashedName: step-34
---

# --description--

`textarea` 元素的行為類似於型別為 `text` 的 `input` 元素，但具有能夠接收多行文字，以及初始文字列數和行數的額外利益。

使用者將能夠註冊時提供個人簡介。為你的 `fieldset` 的結尾添加一個帶有文字 `Provide a bio:` 的 `label`。在 `label` 元素內添加一個 `textarea` 元素。請注意，`textarea` 需要一個結束標籤。

# --hints--

你應該在第三個 `fieldset` 的末尾、現有的 `label` 元素之後添加一個 `label` 元素。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(4)'));
```

你應該給 `label` 一個 `Provide a bio:` 的文字。

```js
assert.match(document.querySelector('fieldset:nth-child(3) > label:nth-child(4)')?.innerText, /Provide a bio/);
```

你應該將 `textarea` 元素巢狀於 `label` 中。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(4) > textarea'));
```

你應該給 `textarea` 元素開啟和關閉標籤。

```js
assert.match(code, /<textarea\s*>[\s\S]*<\/textarea\s*>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset>
      <label>Upload a profile picture: <input type="file" /></label>
      <label>Input your age (years): <input type="number" min="13" max="120" /></label>
      <label>How did you hear about us?
        <select>
          <option value="">(select one)</option>
          <option value="1">freeCodeCamp News</option>
          <option value="2">freeCodeCamp YouTube Channel</option>
          <option value="3">freeCodeCamp Forum</option>
          <option value="4">Other</option>
        </select>
      </label>
--fcc-editable-region--
      
--fcc-editable-region--
    </fieldset>
    <a href="https://www.freecodecamp.org/news/terms-of-service">Read our terms and conditions</a>
    <label for="terms-and-conditions">
      <input id="terms-and-conditions" type="checkbox" required />
      I accept the terms and conditions
    </label>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
