---
id: 60fac4095512d3066053d73c
title: 第 30 步
challengeType: 0
dashedName: step-30
---

# --description--

為表單添加下拉選單很簡單，只需使用 `select` 元素。`select` 元素是 `option` 元素群組的容器，而 `option` 元素則作為每個下拉選單選項的標籤。這兩個元素都需要閉合標籤。

首先在最後一個 `label` 元素下方添加一個 `select` 元素。然後在 `select` 元素內巢狀 5 個 `option` 元素。

# --hints--

你應該為你的第三個 `fieldset` 添加一個 `select` 元素。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > select'));
```

你應該在 `select` 元素內巢狀 5 個 `option` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-child(3) > select > option'), 5);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset>
      <label>Upload a profile picture: <input type="file" /></label>
      <label>Input your age (years): <input type="number" min="13" max="120" /></label>
--fcc-editable-region--
      
--fcc-editable-region--
    </fieldset>
    <a href="https://www.freecodecamp.org/news/terms-of-service">Read our terms and conditions</a>
    <label for="terms-and-conditions">
      <input id="terms-and-conditions" type="checkbox" required />
      I accept the terms and conditions
    </label>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
