---
id: 60fab4a123ce4b04526b082b
title: 第 24 步
challengeType: 0
dashedName: step-24
---

# --description--

目前使用者可以在未勾選單選按鈕輸入項目的情況下提交表單。雖然你之前使用過 `required` 屬性來顯示輸入項目為必填，但在這種情況下不會生效，因為將 `required` 添加到兩個輸入項目會向使用者傳達錯誤的資訊。

為了解決此問題，你可以在第二個 `fieldset` 中的 `label` 元素之前，為你的需求提供背景關係，方法是添加一個帶有文字 `Account type (required)` 的 `legend` 元素。然後，將 `checked` 屬性添加到 `Personal` 輸入，以確保表單提交時包含所需的資料。

# --hints--

你的 `input` 元素應保持 `type` 為 `radio`。

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-child(2) input[type="radio"]'), 2);
```

你應該在第二個 `fieldset` 元素內添加一個 `legend` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(2) > legend'), 1);
```

你應該為你的 `legend` 元素添加 `Account type (required)` 文字。

```js
assert.equal(document.querySelector('fieldset:nth-of-type(2) > legend')?.textContent, 'Account type (required)');
```

你應該為 `Personal` 輸入給予一個 `checked` 屬性。

```js
assert.isTrue(document.querySelector('fieldset:nth-child(2) input[type="radio"]')?.hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
--fcc-editable-region--
      
--fcc-editable-region--
      <label><input type="radio" name="account-type" value="personal" /> Personal</label>
      <label><input type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
