---
id: 62ff8b9dab5ac88e4d3d43a3
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

按照可訪問性最佳實踐，使用 `for` 屬性將 `input` 元素和 `label` 元素關聯在一起。

使用 `first-name`、`last-name`、`email` 和 `new-password` 作爲相應的 `id` 屬性的值。

# --hints--

第一個 `input` 元素應該有一個 `first-name` `id`。

```js
assert.isTrue(document.querySelectorAll('input')?.[0]?.matches('#first-name'))
```

第二個 `input` 元素應該有一個 `last-name` `id`。

```js
assert.isTrue(document.querySelectorAll('input')?.[1]?.matches('#last-name'))
```

第三個 `input` 元素應該有一個 `email` `id`。

```js
assert.isTrue(document.querySelectorAll('input')?.[2]?.matches('#email'))
```

第四個 `input` 元素應該有一個 `new-password` `id`。

```js
assert.isTrue(document.querySelectorAll('input')?.[3]?.matches('#new-password'))
```

第一個 `label` 元素應該有一個 `for` 屬性，值爲 `first-name`。

```js
assert.isTrue(document.querySelectorAll('label')?.[0]?.matches('label[for="first-name"]'))
```

第二個 `label` 元素應該有一個 `for` 屬性，值爲 `last-name`。

```js
assert.isTrue(document.querySelectorAll('label')?.[1]?.matches('label[for="last-name"]'))
```

第三個 `label` 元素應該有一個 `for` 屬性，值爲 `email`。

```js
assert.isTrue(document.querySelectorAll('label')?.[2]?.matches('label[for="email"]'))
```

第四個 `label` 元素應該有一個 `for` 屬性，值爲 `new-password`。

```js
assert.isTrue(document.querySelectorAll('label')?.[3]?.matches('label[for="new-password"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label>Enter Your First Name: <input /></label>
      <label>Enter Your Last Name: <input /></label>
      <label>Enter Your Email: <input /></label>
      <label>Create a New Password: <input /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
