---
id: 63541ef4f96cd82e8e6c788a
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

`method` 屬性指定如何將表單資料傳送到 `action` 屬性中指定的 URL。表單資料可以透過 `GET` 請求作為 URL 參數傳送（使用 `method="get"`），或透過 `POST` 請求作為請求主體中的資料傳送（使用 `method="post"`）。

將 `method` 屬性設定為透過 `POST` 請求傳送你的表單資料。

**注意**：你將在接下來的課程和工作坊中學習更多關於 `GET` 和 `POST` 方法的內容。

# --hints--

你不應該添加新的 `form` 元素。

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

你的 `form` 元素應該有一個 `method` 屬性。

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

你的 `method` 屬性應該設定為 `post`。

```js
assert.match(document.querySelector('form')?.getAttribute('method'), /^post$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
--fcc-editable-region--
  <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
