---
id: 5f332a88dc25a0fd25c7687a
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

咖啡廳的名稱是 `CAMPER CAFE`。在你的 `main` 元素中添加一個 `h1` 元素。將咖啡廳的名稱以大寫字母呈現，使其突出。

# --hints--

你應該有一個開啟的 `<h1>` 頁籤。

```js
assert(code.match(/<h1>/i));
```

你應該有一個結束的 `</h1>` 頁籤。

```js
assert(code.match(/<\/h1>/i));
```

你不應該更改你現有的 `main` 元素。

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

你的 `h1` 元素應該巢狀在你的 `main` 元素中。

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

你的 `h1` 元素應該有文字 `CAMPER CAFE`。

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

