---
id: 5f3326b143638ee1a09ff1e3
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

為了準備創建一些實際內容，請在 `head` 元素下方添加一個 `body` 元素。

# --hints--

你應該有一個開啟的 `<body>` 頁籤。

```js
assert(code.match(/<body>/i));
```

你應該有一個結束的 `</body>` 頁籤。

```js
assert(code.match(/<\/body>/i));
```

你不應該更改你的 `head` 元素。請確保你沒有刪除你的結束標籤。

```js
assert(code.match(/<head>/i));
assert(code.match(/<\/head>/i));
```

你的 `body` 元素應該放在你的 `head` 元素之後。

```js
assert(code.match(/<\/head>[.\n\s]*<body>/im));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```

