---
id: 5f33071498eb2472b87ddee4
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在這個專案中，你將透過建置一個咖啡廳選單來學習 CSS（階層樣式表）的基礎知識。CSS 是用來為 HTML 文件設計樣式的語言。它描述了 HTML 元素應該如何在螢幕上顯示。

正如你在貓咪照片應用程式的最後幾個步驟中學到的，開始建置你的網頁需要一個基本的結構。每個 HTML 文件都應該有一個 `DOCTYPE` 宣告式和 `html` 元素。`DOCTYPE` 告訴瀏覽器該文件使用的是哪個版本的 HTML。而 `html` 元素表現的是包含所有其他元素的根元素。

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

為你的文件添加 `<!DOCTYPE html>` 頁籤，以及一個具有 `lang` 屬性且值為 `en` 的 `html` 元素。

# --hints--

你應該有 `<!DOCTYPE html>` 宣告。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

你應該有一個開頭為 `<html lang="en">` 且語言設定為英文。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你應該有一個結束的 `</html>` 頁籤。

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
