---
id: 5f331e55dfab7a896e53c3a1
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

`title` 是多個元素之一，提供在網頁上看不到的額外資訊，但對搜尋引擎或頁面顯示方式很有用。

在 `head` 元素內，巢狀一個 `meta` 元素，其屬性名為 `charset`，並將其值設為 `utf-8`，以告訴瀏覽器如何為頁面編碼字元。

請記住，`meta` 元素是無結束標籤的空元素。

# --hints--

你應該有一個 `meta` 頁籤。

```js
assert.match(code, /<meta.*?\/?>/is);
```

`meta` 元素是空元素，不應該有結束標籤 `</meta>`。

```js
assert.notMatch(code, /<\/meta>/i);
```

你的 `meta` 頁籤應該有一個 `charset` 屬性。

```js
assert.match(code, /<meta\s+charset\s*/i);
```

你的 `charset` 屬性應該有一個 `utf-8` 的值。

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
