---
id: 5f3313e74582ad9d063e3a38
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `html` 元素內添加一個 `head` 元素，這樣你就可以添加一個 `title` 元素。`title` 元素的文字應該是 `Cafe Menu`。

# --hints--

你應該有一個開啟的 `<head>` 頁籤。

```js
assert(code.match(/<head>/i));
```

你應該有一個結束的 `</head>` 頁籤。

```js
assert(code.match(/<head>/i));
```

你應該有一個開啟的 `<title>` 頁籤。

```js
assert(code.match(/<title>/i));
```

你應該有一個結束的 `</title>` 頁籤。

```js
assert(code.match(/<\/title>/i));
```

你的 `title` 元素應該巢狀在你的 `head` 元素中。

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

你的 `title` 元素應該有文字 `Cafe Menu`。你可能需要檢查你的拼字。

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
