---
id: 5f33294a6af5e9188dbdb8f3
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

是時候為選單內容添加一些內容了。在現有的 `body` 元素中添加一個 `main` 元素。它最終將包含咖啡館提供的咖啡和甜點的價格資訊。

# --hints--

你的程式碼應該有一個開啟的 `<main>` 頁籤。

```js
assert(code.match(/<main>/i));
```

你的程式碼應該有一個結束的 `</main>` 頁籤。

```js
assert(code.match(/<\/main>/i));
```

你不應該更改你的 `body` 元素。請確保你不會意外刪除你的結束標籤。

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

你的 `main` 頁籤應該位於你的 `body` 頁籤內。

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
