---
id: 5f3477ae8466a9a3d2cc953c
title: 第 15 步
challengeType: 0
dashedName: step-15
---

# --description--

既然你已經將 CSS 放在 `styles.css` 檔案中，現在可以刪除 `style` 元素及其所有內容。刪除後，原本置中的文字將會回到左側對齊。

# --hints--

你的程式碼中不應該有任何 `style` 頁籤。

```js
assert(!code.match(/style/i));
```

你的 HTML 檔案中不應該有任何 CSS 選擇器。

```js
assert(!editableContents.includes('style'));
assert(!editableContents.includes('text-align'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```

