---
id: 5f3477cbcb6ba47918c1da92
title: 第 17 步
challengeType: 0
dashedName: step-17
---

# --description--

為了讓頁面的樣式在行動裝置上看起來與桌面或筆記型電腦相似，你需要添加一個帶有特殊 `content` 屬性的 `meta` 元素。

將以下內容添加到 `head` 元素中：

```html
<meta name="viewport" content="width=device-width, initial-scale=1.0" />
```

# --hints--

你的程式碼應該有兩個 `meta` 元素。

```js
assert.lengthOf(code.match(/<meta.*?\/?>/g),2);
```

你的 `meta` 元素應該有一個 `name` 屬性，且其值為 `viewport`。

```js
const meta = document.querySelectorAll('meta');
assert.exists(meta[0]?.outerHTML?.match(/name=('|")viewport\1/) || meta[1]?.outerHTML?.match(/name=('|")viewport\1/));
```

你的 `meta` 元素應該有一個 `content` 屬性，其值為 `width=device-width, initial-scale=1.0`。

```js
const meta = document.querySelectorAll('meta');
assert.exists(meta[0]?.outerHTML?.match(/content=('|")width=device-width, initial-scale=1.0\1/) || meta[1]?.outerHTML?.match(/content=('|")width=device-width, initial-scale=1.0\1/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
