---
id: 5f344f9c805cd193c33d829c
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

你可以透過在 `style` 元素中指定並為它設定屬性，來為元素添加樣式，如下所示：

```css
element {
 property: value;
}
```

將 `h1` 元素的內容置中，方法是將其 `text-align` 屬性設定為值 `center`。

# --hints--

你應該在你的 `style` 元素中有一個 `h1` 選擇器。

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasSelector);
```

你的 `text-align` 屬性應該有 `center` 的值。

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert(hasTextAlign);
```

你的 `h1` 選擇器應該將 `text-align` 屬性設為 `center`。

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

