---
id: 5f3477cb303c5cb61b43aa9b
title: 步驟 18
challengeType: 0
dashedName: step-18
---

# --description--

文字再次置中，因此指向 CSS 檔案的聯結有效。為你的檔案添加另一個樣式，將 `body` 元素的 `background-color` 屬性更改為 `brown`。

# --hints--

你應該使用 `body` 選擇器。

```js
const hasBody = new __helpers.CSSHelp(document).getStyle('body');
assert(hasBody);
```

你應該將 `background-color` 屬性設為 `brown`。

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'brown');
assert(hasBackground);
```

你的 `body` 元素應該有 `brown` 背景。

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-color');
assert(bodyBackground === 'brown');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--
h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```

