---
id: 5f344fc1520b6719f2e35605
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

選單上會有兩個區塊，一個是咖啡區塊，一個是甜點區塊。為你的 `main` 元素添加一個 `section` 元素，這樣你就有地方放所有可用的咖啡。

# --hints--

你應該有一個開啟的 `<section>` 頁籤。

```js
assert(code.match(/<section\s*>/i));
```

你應該有一個結束的 `</section>` 頁籤。

```js
assert(code.match(/<\/section\s*>/i));
```

你不應該更改你現有的 `main` 元素。請確保你沒有刪除結束標籤。

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

你的 `section` 元素應該位於你的 `main` 元素內。

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert(section.parentElement === main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
