---
id: 5f3477cb2e27333b1ab2b955
title: 第 16 步
challengeType: 0
dashedName: step-16
---

# --description--

現在你需要聯結 `styles.css` 檔案，這樣樣式才會再次套用。在 `head` 元素內，添加一個 `link` 元素。為它設定一個 `rel` 屬性，值為 `"stylesheet"`，以及一個 `href` 屬性，值為 `"styles.css"`。

請注意，`link` 元素是空元素。

# --hints--

你的程式碼應該有一個 `link` 元素。

```js
const link = document.querySelector('link');
assert.isNotNull(link);
```

`link` 元素是空元素，不應該有結束標籤 `</link>`。

```js
assert.notMatch(code, /<\/link>/i);
```

你不應該更改你現有的 `head` 元素。請確保你沒有刪除結束標籤。

```js
const headElementCount = document.querySelectorAll('head')?.length;
assert.strictEqual(headElementCount, 1);
```

你應該有一個 `link` 元素。

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 1);
```

你的 `link` 元素應該位於你的 `head` 元素內。

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

你的 `link` 元素應該有一個 `rel` 屬性，其值為 `stylesheet`。

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

你的 `link` 元素應該有一個 `href` 屬性，其值為 `styles.css`。

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
