---
id: 5f3477ae34c1239cafe128be
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

你現在有三個具有相同樣式的型別選擇器。你可以透過建立一個選擇器列表，將相同的樣式群組添加到多個元素。每個選擇器之間用逗號分隔，如下所示：

```css
selector1, selector2 {
  property: value;
}
```

刪除三個現有的型別選擇器，並以一個選擇器列表取代，將 `h1`、`h2` 和 `p` 元素的文字置中。

# --hints--

你應該為所有三個元素 `h1, h2, p` 使用單一型別選擇器。請確保使用該順序。

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert(hasSelector);
```

你的 `style` 元素中應該只有一個選擇器。

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert(selectors.length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

