---
id: 5f3477aefa51bfc29327200b
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

你已經透過在 `style` 頁籤內撰寫 CSS 來為三個元素設計樣式。這樣可以運作，但由於樣式會更多，最好將所有樣式放在獨立的檔案中並連結到它。

已為你創建了獨立的 `styles.css` 檔案。你可以透過編輯器頂部的分頁在檔案之間切換。

首先將你已經創建的樣式重寫到 `styles.css` 檔案中。請確保你排除開啟和關閉的 `style` 頁籤。

# --hints--

你的 `styles.css` 檔案應該包含 `h1, h2, p` 型別選擇器。

```js
assert(editableContents.replace(/[\s\n]*/g, "").match(/(h1|h2|p),(h1|h2|p),(h1|h2|p){/));
```
 
你的選擇器應該將 `text-align` 屬性設為 `center`。

```js
assert(editableContents.match(/text-align:\s*center;?/));
```

你的程式碼不應包含 `<style>` 頁籤。

```js
assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
