---
id: 5f356ed69db0a491745e2bb6
title: 步驟 27
challengeType: 0
dashedName: step-27
---

# --description--

要將類別的樣式套用到 `div` 元素，請移除 `id` 屬性，並在 `div` 元素的開啟標籤中添加 `class` 屬性。請確保將類別值設定為 `menu`。

# --hints--

你的 `div` 應該仍然會呈現。請確定你沒有錯誤地撰寫 `<div>` 頁籤。

```js
assert.lengthOf(document.querySelectorAll('div'),1);
```

你的 `div` 元素應該有 `menu` 類別。

```js
assert.isTrue(document.querySelector('div').classList.contains('menu'));
```

你的 `div` 元素不應該再有 `menu` id。

```js
assert.lengthOf(document.querySelectorAll('div#menu'),0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
--fcc-editable-region--
    <div id="menu">
--fcc-editable-region--
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

