---
id: 5f356ed60a5decd94ab66986
title: 第 22 步
challengeType: 0
dashedName: step-22
---

# --description--

CSS 中的註解長這樣：

```css
/* comment here */
```

在你的樣式表中，將包含 `background-color` 屬性和值的那一行註解掉，這樣你就可以看到只對 `#menu` 元素進行樣式設定的效果。這將使背景再次變成白色。

# --hints--

你應該將 CSS 中的 `background-color: burlywood;` 行註解掉。

```js
assert(code.match(/\/\*\s*background-color:\s*burlywood;?\s*\*\//i));
```


你的 `body` 應該有白色背景。

```js
const body = document.querySelector("body");
const compStyles = window.getComputedStyle(body);
const backgroundColor = compStyles?.getPropertyValue('background-color'); 
assert.equal(backgroundColor,"rgba(0, 0, 0, 0)")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
--fcc-editable-region--
  background-color: burlywood;
--fcc-editable-region--
}

h1, h2, p {
  text-align: center;
}

#menu {
  width: 300px;
}
```

