---
id: 5f356ed60785e1f3e9850b6e
title: 第 24 步
challengeType: 0
dashedName: step-24
---

# --description--

現在很容易看出文字置中在 `#menu` 元素內。目前，`#menu` 元素的寬度是以像素（`px`）指定的。 

將 `width` 屬性的值更改為 `80%`，使其寬度為其父元素（`body`）的 80%。

# --hints--

你應該將 `width` 屬性設為 `80%`。

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '80%');
assert(hasWidth);
```

你不應該有 `width` 屬性為 `300px`。

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(!hasWidth);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--
#menu {
  width: 300px;
  background-color: burlywood;
}
--fcc-editable-region--
```

