---
id: 5f356ed63c7807a4f1e6d054
title: 第 21 步
challengeType: 0
dashedName: step-21
---

# --description--

現在的目標是讓 `div` 不佔據整個頁面的寬度。CSS 的 `width` 屬性非常適合用來達成這個目的。

你可以使用 `id` 選擇器來標的具有 `id` 屬性的特定元素。`id` 選擇器是透過將井字號符號 `#` 直接放在元素的 `id` 值前面來定義的。<dfn>例如</dfn>，如果一個元素的 `id` 是 `cat`，那麼你可以這樣標的該元素：

```css 
#cat {
  width: 250px;
}
```

使用 `#menu` 選擇器為你的元素設定寬度為 `300px`。

# --hints--

你應該有一個 `#menu` 選取器。

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle("#menu");
assert(hasDiv);
```

你應該將 `width` 屬性設為 `300px`。

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(hasWidth);
```

你的 `div` 應該有 300px 的寬度。

```js
const divWidth = new __helpers.CSSHelp(document).getStyle("#menu")?.getPropertyValue('width');
assert(divWidth === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: burlywood;
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```

