---
id: 5f3c866d5414453fc2d7b480
title: 步驟 31
challengeType: 0
dashedName: step-31
---

# --description--

在現有的 coffee/price 對組下方，使用包含兩個巢狀 `p` 元素的 `article` 元素，添加以下咖啡和價格：

```md
Caramel Macchiato 3.75
Pumpkin Spice 3.50
Hazelnut 4.00
Mocha 4.50
```

如同之前，第一個 `p` 元素的文字應該包含咖啡風味，第二個 `p` 元素的文字應該包含價格。

# --hints--

你應該有五個 `article` 元素。

```js
assert.lengthOf(document.querySelectorAll('article'),5);
```

每個 `article` 元素應該有兩個 `p` 元素。

```js
const articles = document.querySelectorAll('article');
assert.lengthOf(articles?.[0]?.children, 2);
assert.lengthOf(articles?.[1]?.children, 2);
assert.lengthOf(articles?.[2]?.children, 2);
assert.lengthOf(articles?.[3]?.children, 2);
assert.lengthOf(articles?.[4]?.children, 2);
```

你的第一個 `article` 元素應該包含帶有文字 `French Vanilla` 和 `3.00` 的 `p` 元素。

```js
const children = document.querySelector('article')?.children;
assert.match(children?.[0]?.innerText.trim(), /French Vanilla/i);
assert.match(children?.[1]?.innerText.trim(), /3\.00/i);
```

你的第二個 `article` 元素應該有包含文字 `Caramel Macchiato` 和 `3.75` 的 `p` 元素。

```js
const children = document.querySelectorAll('article')?.[1]?.children;
assert.match(children?.[0]?.innerText.trim(), /Caramel Macchiato/i);
assert.match(children?.[1]?.innerText.trim(), /3\.75/i);
```

你的第三個 `article` 元素應該有包含文字 `Pumpkin Spice` 和 `3.50` 的 `p` 元素。

```js
const children = document.querySelectorAll('article')?.[2]?.children;
assert.match(children?.[0]?.innerText.trim(), /Pumpkin Spice/i);
assert.match(children?.[1]?.innerText.trim(), /3\.50/i);
```

你的第四個 `article` 元素應該有包含文字 `Hazelnut` 和 `4.00` 的 `p` 元素。

```js
const children = document.querySelectorAll('article')?.[3]?.children;
assert.match(children?.[0]?.innerText.trim(), /Hazelnut/i);
assert.match(children?.[1]?.innerText.trim(), /4\.00/i);
```

你的第五個 `article` 元素應該有包含文字 `Mocha` 和 `4.50` 的 `p` 元素。

```js
const children = document.querySelectorAll('article')?.[4]?.children;
assert.match(children?.[0]?.innerText.trim(), /Mocha/i);
assert.match(children?.[1]?.innerText.trim(), /4\.50/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
            <p>French Vanilla</p>
            <p>3.00</p>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
