---
id: 5f7b87422a560036fd03ccff
title: 步驟 50
challengeType: 0
dashedName: step-50
---

# --description--

為你剛剛添加的兩個 `p` 元素，將 `dessert` 添加為第一個 `p` 元素的 `class` 屬性值，並將 `price` 添加為第二個 `p` 元素的 `class` 屬性值。

# --hints--

你應該有一個帶有 `dessert` 類別的 `p` 元素。

```js
assert.lengthOf(document.querySelectorAll('.dessert'), 1);
```

你的 `p` 元素中帶有 `Donut` 文字應該有 `dessert` 類別。

```js
assert.match(document.querySelector('.dessert')?.innerText, /donut/i);
```

你的 `p` 元素中帶有 `1.50` 文字，應該有 `price` 類別。

```js
const lastElement = [...document.querySelectorAll('.price')].at(-1);
assert.match(lastElement?.textContent, /1\.50/);
```

你的 `p` 元素之間不應該有任何空格。

```js
assert(!code.match(/<\/p>\s+<p/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
          <article class="item">
--fcc-editable-region--
            <p>Donut</p><p>1.50</p>
--fcc-editable-region--
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%;
}
```

