---
id: 6133acc353338c0bba9cb553
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

`head` の中に関して最後に、`title` 要素はスクリーンリーダーでページの内容を理解するのに役立ちます。 しかも、それは _SEO_ にとって重要な部分です。

Give your page a `title` that is descriptive and concise.

# --hints--

You should add a `title` element to the `head`.

```js
assert.exists(document.querySelector('head > title'));
```

You should not make the `title` longer than 60 characters.

```js
assert.isAtMost(document.querySelector('head > title')?.textContent?.length, 60);
```

Try being more descriptive with your `title` element. _Hint: At least 15 characters_

```js
assert.isAtLeast(document.querySelector('head > title')?.textContent?.length, 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
