---
id: 614ccc21ea91ef1736b9b578
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

For this workshop, you have been provided with a basic boilerplate.

`html` 要素に `lang` 属性を加えるところから、このアクセシビリティの旅を始めましょう。 これはスクリーンリーダーがページの言語を特定するのに使用されます。

Also, remember to link your stylesheet to the page.

# --hints--

You should give the `html` element a `lang` attribute. _Hint: You can use the value `en` for English._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

You should have a `link` element in your `head` element.

```js
assert.exists(document.querySelector("head>link"));
```

Your `link` element should have `rel="stylesheet"` and `href="styles.css"`.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>

  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
