---
id: 69446f18fdcb11aa516f7c5c
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستنشئ دالة `countup` التي تُنشئ مصفوفة تحتوي أرقامًا من 1 إلى رقم معين باستخدام العودية.

ابدأ بإنشاء دالة فارغة باسم `countup` تأخذ معلمة واحدة.

# --hints--

يجب أن يكون لديك دالة باسم `countup`.

```js
assert.isFunction(countup);
```

يجب أن تقبل دالة `countup` معلمة واحدة.

```js
assert.strictEqual(countup.length, 1);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
