---
id: 69447054012a22f38c944ff0
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

أضف جملة `if` داخل دالتك `countup`.

تحقق مما إذا كان `number` أقل من `1` وأرجع مصفوفة فارغة.

سيكون هذا هو الحالة الأساسية لاستدعاء الدالة ذاتيًا.

# --hints--

عندما يكون `number` أقل من `1` يجب أن تُرجع الدالة مصفوفة فارغة.

```js
assert.deepStrictEqual(countup(0), []);
```

عندما لا يكون `number` أقل من `1` يجب ألا تُرجع الدالة أي شيء.

```js
assert.isUndefined(countup(1)); 
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
