---
id: 6948705160b66cf93a916e29
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

أعلن عن مصفوفة فارغة باسم `countArray` داخل الدالة `countup` مباشرة فوق جملة `if` الخاصة بك.

ستخزن هذه المصفوفة قيمة `number` عند كل استدعاء للدالة.

# --hints--

يجب أن تعلن عن مصفوفة فارغة باسم `countArray`.

```js
assert.match(countup.toString(), /(let|const|var)\s+countArray\s*=\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region-- 
  
--fcc-editable-region--
  if (number < 1) {
    return [];
  }
}
```
