---
id: 694870bcf3c6874112b6f951
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

لكل الأعداد الأكبر من أو تساوي `1`، أنشئ كتلة `else` وداخلها، استدعِ الدالة `countup` مع `number - 1` وعيّن النتيجة إلى `countArray`.

ثم استخدم `.push()` لإضافة `number` إلى `countArray`، وأخيرًا أرجع `countArray`.

هذا يضمن أن الاستدعاء الذاتي يصل في النهاية إلى الحالة الأساسية وينهي التنفيذ.

# --hints--

يجب تعيين نتيجة استدعاء `countup` مع `number - 1` إلى `countArray`.

```js
assert.match(countup.toString(), /countArray\s*=\s*countup\(\s*number\s*-\s*1\s*\)/);
```

يجب دفع `number` إلى `countArray`.

```js
assert.match(countup.toString(), /countArray\.push\(\s*number\s*\)/);
```

يجب إرجاع `countArray`.

```js
assert.deepStrictEqual(countup(3), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
--fcc-editable-region--
  } 
--fcc-editable-region--
}
```
