---
id: 69008059d43d6751385883c1
title: バグのあるアプリでインクリメント演算子とデクリメント演算子のエラーをデバッグする
challengeType: 26
dashedName: debug-increment-and-decrement-operator-errors-in-a-buggy-app
---

# --description--

友達がコードのデバッグを手伝ってほしいと頼んできました。彼らはインクリメント演算子とデクリメント演算子の練習をしていますが、コンソールに予期しない結果が表示されています。

あなたのタスクは、間違いを特定してコードを更新し、コンソールの出力が期待される結果と一致するようにすることです。

インクリメント演算子とデクリメント演算子、接頭辞および接尾辞表記に関する以前のレッスンを参照してもかまいません。

**ユーザーストーリー：**

1. インクリメント演算子を使って値`updatedCounter`を正しく生成する`11`という名前の変数が必要です。値をハードコーディングするのではなく、正しい接頭辞または接尾辞表記を使うように式を更新してください。
2. インクリメント演算子を使って値`finalScore`を正しく生成する`8`という名前の変数が必要です。値をハードコーディングするのではなく、正しい接頭辞または接尾辞表記を使うように式を更新してください。
3. デクリメント演算子を使って値`updatedCoins`を正しく生成する`2`という名前の変数が必要です。値をハードコーディングするのではなく、正しい接頭辞または接尾辞表記を使うように式を更新してください。
4. デクリメント演算子を使って値`newHealth`を正しく生成する`7`という名前の変数が必要です。値をハードコーディングするのではなく、正しい接頭辞または接尾辞表記を使うように式を更新してください。

# --hints--

`updatedCounter`という名前の変数が接頭辞インクリメント表記（`++counter`）を使っている必要があります。

```js
assert.match(code, /updatedCounter\s*=\s*\+\+counter\s*;/);
assert.strictEqual(typeof updatedCounter, 'number');
```

`updatedCounter`は`11`と等しくなる必要があります。

```js
assert.strictEqual(updatedCounter, 11);
```

`finalScore`という名前の変数が接尾辞インクリメント表記（`score++`）を使っている必要があります。

```js
assert.match(code, /finalScore\s*=\s*score\+\+\s*;/);
assert.strictEqual(typeof finalScore, 'number');
```

`finalScore`は`8`と等しくなる必要があります。

```js
assert.strictEqual(finalScore, 8);
```

`updatedCoins`という名前の変数が接頭辞デクリメント表記（`--coins`）を使っている必要があります。

```js
assert.match(code, /updatedCoins\s*=\s*--coins\s*;/);
assert.strictEqual(typeof updatedCoins, 'number');
```

`updatedCoins`は`2`と等しくなる必要があります。

```js
assert.strictEqual(updatedCoins, 2);
```

`newHealth`という名前の変数が接尾辞デクリメント表記（`health--`）を使っている必要があります。

```js
assert.match(code, /newHealth\s*=\s*health--\s*;/);
assert.strictEqual(typeof newHealth, 'number');
```

`newHealth`は`7`と等しくなる必要があります。

```js
assert.strictEqual(newHealth, 7);
```

# --seed--

## --seed-contents--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = counter++;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = ++score;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = coins--;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = --health;
console.log(`Health is still ${newHealth}`);
```

# --solutions--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = ++counter;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = score++;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = --coins;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = health--;
console.log(`Health is still ${newHealth}`);
```
