---
id: 6690e10ebe2181212abc9652
title: Diseña una tarjeta de presentación
challengeType: 25
dashedName: design-a-business-card
demoType: onClick
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Deberías tener un `div` con un atributo `class` con un valor de `business-card` que contenga todos los otros elementos.
1. Dentro del elemento `.business-card`, debe haber un elemento `img` con un atributo `class` con un valor de `profile-image`. Puedes establecer la fuente de la imagen en `https://cdn.freecodecamp.org/curriculum/labs/flower.jpg` si lo deseas. Debe haber un `alt` con una descripción significativa.
1. Dentro del elemento `.business-card`, después del elemento `img`, debería haber tres elementos `p` con un atributo `class` que tenga, respectivamente, un valor de `full-name`, `designation`, y `company`.
1. El primer elemento `p` debe contener tu nombre.
1. El segundo elemento `p` debe contener tu designación.
1. El tercer elemento `p` debe tener el nombre de tu empresa.
1. Debe haber un elemento `hr` debajo del tercer elemento `p`.
1. Después del elemento `hr`, debe haber dos elementos `p`. En el primer elemento `p`, deberías tener una dirección de correo electrónico como texto. En el segundo elemento `p`, el texto debe ser un número de teléfono.
1. Después de los dos elementos `p`, debe haber un elemento `a` con una clase `portfolio-link`, el texto `Portfolio` y debe enlazar a una URL válida.
1. Debe haber un elemento `hr` después del elemento `a` que contiene el enlace del portafolio.
1. Deberías tener otro elemento `div` con un atributo `class` con un valor de `social-media`. Dentro de este elemento, debe haber un elemento `h2` con el texto `Connect with me`.
1. Dentro del os elementos `.social-media`, debe haber tres elementos `a`. Los elementos `a` deben tener el texto `Twitter`, `LinkedIn` y `GitHub` respectivamente con enlaces a sitios web válidos.
1. Deberías enlazar la hoja de estilo `styles.css` usando la etiqueta `link` en la sección `head` del archivo HTML.
1. El fondo de tu página debería ser `rosybrown`. La fuente general debe ser `Arial`, con una alternativa de `sans-serif`.
1. El selector `.business-card` debería tener propiedades para establecer el elemento con un ancho de `300px` y un color de fondo a tu elección. También deberías establecer un `padding` de `20px` alrededor y un margen superior de `100px`. El texto debería estar alineado al centro y el tamaño de la fuente debería ser `16px`. Los márgenes izquierdo y derecho deben establecerse en `auto`.
1. El selector `.profile-image` debería tener una propiedad `max-width` con un valor de `100%`.
1. Todos los elementos `p` deben tener un margen superior e inferior de `5px`.
1. Todos los enlaces en la página no deberían estar subrayados.

# --hints--

Deberías tener un `div` principal con un atributo `class` con un valor de `business-card`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'business-card');
```

Dentro del elemento `.business-card`, debería haber un elemento `img` con un atributo `class` con un valor de< code>profile-image</code>.

```js
assert.exists(document.querySelector('img.profile-image'));
```

La fuente de la imagen debe establecerse en una imagen válida.

```js
assert.isAbove(document.querySelector('.profile-image')?.getAttribute('src')?.length, 0);
```

El atributo `alt` de la imagen debe establecerse en una descripción significativa.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

Dentro del elemento `.business-card`, después del elemento `img`, debe haber un elemento `p` con un atributo `class` con un valor de `full-name`.

```js
assert.exists(document.querySelector('div.business-card p.full-name'));
```

Después del elemento `.full-name`, debe haber un elemento `p` con un atributo `class` con un valor de< code>designation</code>.

```js
assert.exists(document.querySelector('div.business-card p.designation'));
```

Después del elemento `.designation`, debe haber un elemento `p` con un atributo `class` con un valor de `company`.

```js
assert.exists(document.querySelector('div.business-card p.company'));
```

El primer elemento `p`, el de clase `full-name`, debe contener tu nombre.

```js
assert.isNotEmpty(document.querySelector(".full-name")?.innerText.trim());
```

El segundo elemento `p`, el atributo `class` de `designation`, debe contener tu designación.

```js
assert.isNotEmpty(document.querySelector(".designation")?.innerText.trim());
```

El tercer elemento `p`, el de atributo `class` con un valor de`company`, debe contener el nombre de tu empresa.

```js
assert.isNotEmpty(document.querySelector(".company")?.innerText.trim());
```

Después del tercer elemento `p`, el de atributo `class` con un valor de `company`, debe haber un elemento `hr`.

```js
assert.exists(document.querySelector('p.company + hr'));
```

Después del primer elemento `hr`, debe haber un elemento `p` con una dirección de correo electrónico como tu texto.

```js
assert.isNotEmpty(document.querySelectorAll('.business-card hr + p')[0].innerText.trim());
```

Después del elemento `p` de correo electrónico, debe haber otro elemento `p` con un número de teléfono como tu texto.

```js
assert.match(document.querySelectorAll(".business-card p")[4].textContent,/.+/);
```

Después del elemento `p` del número de teléfono, debe haber un elemento `a` con la clase `portfolio-link`. 

```js
assert.exists(document.querySelector('a.portfolio-link'));
```

El elemento `.portfolio-link` debe tener el texto `Portfolio`.

```js
assert.equal(document.querySelector('a.portfolio-link').innerText.trim(), 'Portfolio');
```

Tu elemento `.portfolio-link` debe tener un `href` válido.

```js
assert.isAbove(document.querySelector('.portfolio-link').getAttribute('href').length, 0);
```

Después del elemento `a` que contiene tu portafolio, debe haber otro elemento `hr` como divisor.

```js
assert.exists(document.querySelector('a + hr'));
```

Después del segundo elemento `hr`, debe haber un elemento `div` con un atributo `class` con un valor de `social-media`.

```js
assert.exists(document.querySelector('div.business-card div.social-media'));
```

Dentro del elemento `.social-media`, debe haber un elemento `h2` con el texto `Connect with me`.

```js
assert.equal(document.querySelector('.social-media h2')?.textContent.trim(), 'Connect with me');
```

Dentro del elemento `.social-media` debe haber tres elementos `a`.

```js
assert.lengthOf(document.querySelectorAll('.social-media a'), 3);
```

El primer elemento `a` debe tener el texto `Twitter`.

```js
assert.equal(document.querySelectorAll(".social-media a")[0]?.textContent.trim(), 'Twitter');
```

El `href` del primer elemento `a` debe apuntar a un enlace válido.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[0]?.getAttribute('href')?.length, 0);
```

El segundo elemento `a` debe tener el texto `LinkedIn`.

```js

assert.equal(document.querySelectorAll(".social-media a")[1]?.textContent.trim(), 'LinkedIn');
```

El `href` del segundo elemento `a` debe apuntar a un enlace válido.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[1]?.getAttribute('href')?.length, 0);
```

El tercer elemento `a` debe tener el texto `GitHub`.

```js
assert.equal(document.querySelectorAll(".social-media a")[2]?.textContent.trim(), 'GitHub');
```

El `href` del tercer elemento `a` debe apuntar a un enlace válido.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[2]?.getAttribute('href')?.length, 0);
```

Su elemento `link` debe estar dentro de su elemento `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Su elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Tu elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');
```

Deberías establecer el color de fondo de la página en el selector `body` como `rosybrown`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background-color'), 'rosybrown');
```

Debes establecer la fuente de la página en `Arial` en el elemento `body` con una alternativa de `sans-serif`.

```js
const el = new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('font-family');
const condition1= el === 'Arial, sans-serif';
const condition2= el === '"Arial", sans-serif';
assert.isTrue(condition1 || condition2);
```

Tu selector `.business-card` debería tener una propiedad `width` con un valor de `300px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('width'), '300px');
```

Tu selector `.business-card` debería tener una propiedad `background-color` con un valor de color válido.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.business-card')?.backgroundColor);
```

Tu selector `.business-card` debería tener una propiedad `padding` con un valor de `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('padding'), '20px');
```

Tu selector `.business-card` debería tener un margen superior de `100px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-top'), '100px');
```

El texto dentro del elemento `.business-card` debería estar alineado al centro.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('text-align'), 'center');
```

El tamaño de la fuente del texto dentro del elemento `.business-card` debería ser `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('font-size'), '16px');
```

Los márgenes izquierdo y derecho del elemento `.business-card` deberían establecerse en `auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-left'), 'auto');
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-right'), 'auto');
```

Tu selector `.profile-image` debería tener una propiedad `max-width` con un valor de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image')?.getPropVal('max-width'), '100%');
```

Tus elementos `p` deben tener un margen superior e inferior de `5px`.

```js
const marginTop = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-top');
const marginBottom = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-bottom');

assert.equal(marginTop, '5px');
assert.equal(marginBottom, '5px');
```

Los enlaces de la página no deberían tener subrayado.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('a')?.getPropVal('text-decoration'), 'none');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Business Card</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Digital Business Card</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/flower.jpg" alt="A vibrant flower in full bloom" class="profile-image">
        <p class="full-name">Your Name</p>
        <p class="designation">Software Developer</p>
        <p class="company" >@freeCodeCamp</p>
        <hr>
        <p>Email: your-email@example.com</p>
        <p>Phone: (123) 456-7890</p>
        <a class="portfolio-link" href="http://example.com">Portfolio</a>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    max-width: 100%;
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    margin: 20px 0;
}
```
