---
id: 66629f407d679d3105e8317f
title: بناء لعبة حجر ورقة مقص
challengeType: 25
dashedName: top-build-a-rock-paper-scissors-game
---

# --description--
ستلعب لعبتك ضد الكمبيوتر. ستكتب دالة تُرجع عشوائيًا `"rock"` أو `"paper"` أو `"scissors"`.

لا تحتاج للقلق بشأن الجزء الأمامي من اللعبة. ستكتب فقط منطق اللعبة. افتح ملف `script.js` وابدأ باتباع التعليمات.

**قصص المستخدم:**

1. يجب أن يكون لديك دالة باسم `getComputerChoice`.

1. يجب أن تُرجع دالتك `getComputerChoice` عشوائيًا `"rock"` أو `"paper"` أو `"scissors"`.

**تلميح:** طريقة [Math.random](https://developer.mozilla.org/en-US/docs/Web/JavaScript/Reference/Global_Objects/Math/random) تُرجع رقمًا عشوائيًا أكبر من أو يساوي 0 وأقل من 1. فكر كيف تستخدم هذا لإرجاع أحد الخيارات بشكل شرطي.

ستلعب لعبتك بواسطة لاعب بشري. ستكتب دالة تأخذ اختيار المستخدم وتُرجعه.

1. أنشئ دالة باسم `getHumanChoice`.

1. اكتب الشفرة بحيث تُرجع `getHumanChoice` أحد الخيارات الصحيحة بناءً على ما يُدخله المستخدم.

**تلميح:** استخدم طريقة [prompt](https://developer.mozilla.org/en-US/docs/Web/API/Window/prompt) للحصول على مُدخَل المستخدم.

ستتابع لعبتك تسجيل نقاط اللاعب. ستكتب متغيرات لتتبع نقاط اللاعب.

1. أنشئ متغيرين جديدين باسم `humanScore` و `computerScore` في النطاق العام.

1. ابدأ تهيئة تلك المتغيرات بالقيمة `0`.

ستُلعب اللعبة جولة بجولة. ستكتب دالة تأخذ اختيارات اللاعب البشري والكمبيوتر كمعلمات، تلعب جولة واحدة، تزيد نقاط الفائز بالجولة، وتُسجل إعلان الفائز.

1. أنشئ دالة جديدة باسم `playRound`.

1. عرّف معلمتين لـ `playRound`. المعلمة الأولى `humanChoice` والثانية `computerChoice`. استخدم هاتين المعلمتين لأخذ اختيارات اللاعب البشري والكمبيوتر كمعلمات.

1. اجعل معلمة `humanChoice` في دالتك غير حساسة لحالة الأحرف حتى يتمكن اللاعبون من إدخال `"rock"` أو `"ROCK"` أو `"RocK"` أو غيرها من التنويعات.

1. اكتب الشفرة لدالتك `playRound` التي تُرجع قيمة نصية تمثل الفائز بالجولة.

- إذا كانت الجولة تعادلًا، يجب أن تُرجع `"It's a tie!"`.

- إذا فاز اللاعب، يجب أن تُرجع `"You win! [player choice] beats [computer choice]"`.

- إذا فاز الكمبيوتر، يجب أن تُرجع `"You lose! [computer choice] beats [player choice]"`.

1. زد قيمة المتغير `humanScore` أو `computerScore` بناءً على الفائز بالجولة.

ستلعب لعبتك 3 جولات. ستكتب دالة باسم `playGame` تستدعي `playRound` لتلعب 3 جولات، تتتبع النقاط، وتعلن الفائز في النهاية.

1. أنشئ دالة جديدة باسم `playGame`.

1. أنشئ حلقة تلعب 3 جولات وتستدعي دالة `playRound` في كل مرة مع اختيارات اللاعب البشري والكمبيوتر كمعلمات.

1. في نهاية اللعبة، أرجع الفائز بناءً على من فاز بأكثر الجولات.

- إذا فاز اللاعب البشري بعدد جولات أكثر من الكمبيوتر، أرجع رسالة تقول `"You win the game!"`.

- إذا فاز الكمبيوتر بعدد جولات أكثر من اللاعب البشري، أرجع رسالة تقول `"You lose the game!"`.

# --hints--

يجب أن يكون لديك دالة تسمى `getComputerChoice`

```js
assert.isFunction(getComputerChoice);
```

يجب أن تُرجع دالتك `getComputerChoice` عشوائيًا `"rock"` أو `"paper"` أو `"scissors"`.

```js
const counts = {}
for (let i = 0; i < 100; i++) {
  const result = getComputerChoice();
  counts[result] = (counts[result] ?? 0) + 1;
}
assert.lengthOf(Object.keys(counts), 3);
assert.include(Object.keys(counts), "rock");
assert.include(Object.keys(counts), "paper"); 
assert.include(Object.keys(counts), "scissors");
```

يجب أن يكون لديك دالة تسمى `getHumanChoice`

```js
assert.isFunction(getHumanChoice);
```

يجب أن يكون لديك متغيران باسم `humanScore` و `computerScore` في النطاق العام.

```js
assert.exists(humanScore);
assert.exists(computerScore);
```

يجب أن يكون لديك دالة تسمى `playRound`

```js
assert.isFunction(playRound);
```

يجب أن تأخذ دالتك `playRound` اختيارات اللاعب البشري والكمبيوتر كمعلمات باسم `humanChoice` و `computerChoice`.

```js
assert.match(playRound.toString(), /\s*(?:\bhumanChoice\b\s*,\s*\bcomputerChoice\b)/);
```

يجب أن تكون دالتك `playRound` غير حساسة لحالة الأحرف حتى يتمكن اللاعبون من إدخال `"rock"` أو `"ROCK"` أو `"RocK"` أو غيرها من التنويعات.

```js
assert.match(playRound.toString(), /\bhumanChoice\s*\.toLowerCase\(\)/);
```

يجب أن تُرجع دالتك `playRound` النص `"It's a tie!"` إذا كانت اختيارات اللاعب البشري والكمبيوتر متطابقة.

```js
assert.equal(playRound("rock", "rock"), "It's a tie!");
```

يجب أن تُرجع دالتك `playRound` النص `"You win! [player choice] beats [computer choice]"` إذا فاز اللاعب.

```js
assert.equal(playRound("rock", "scissors"), "You win! rock beats scissors");
```

يجب أن تُرجع دالتك `playRound` النص `"You lose! [computer choice] beats [player choice]"` إذا فاز الكمبيوتر.

```js
assert.equal(playRound("rock", "paper"), "You lose! paper beats rock");
```

يجب أن تزيد دالتك `playRound` قيمة المتغير `humanScore` أو `computerScore` بناءً على الفائز بالجولة.

```js
humanScore = 0;
computerScore = 0;
playRound("rock", "scissors");
assert.equal(humanScore, 1);
```

يجب أن يكون لديك دالة تسمى `playGame`.

```js
assert.isFunction(playGame);
```

يجب أن تستخدم حلقة للعب 3 جولات.

```js
assert.match(playGame.toString(), /\bfor\s*\(/);
```

يجب أن تُرجع الفائز في اللعبة بناءً على من فاز بأكثر الجولات.

```js
window.prompt = () => "rock";
assert.match(playGame(), /You (win|lose) the game!/);
```

# --seed--

## --seed-contents--

```js
const hand = ['rock', 'paper', 'scissors'];
```

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

# --solutions--

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

```js
const hand = ['rock', 'paper', 'scissors'];


let computerScore = 0;
let humanScore = 0;

const getComputerChoice = () => {
  return hand[Math.floor(Math.random() * hand.length)];
}

const getHumanChoice = () => {
  return prompt();
}

const playRound = (humanChoice, computerChoice) => {
  humanChoice = humanChoice.toLowerCase()

  const tie = "It's a tie!"
  const win = `You win! ${humanChoice} beats ${computerChoice}`
  const lose = `You lose! ${computerChoice} beats ${humanChoice}`


  if (humanChoice == 'rock') {
    if (computerChoice == 'rock') {
      computerScore++;
      humanScore++;
      return tie;
    } else if (computerChoice == 'paper') {
      computerScore++;
      return lose;
    } else {
      humanScore++;
      return win;
    }
  } else if (humanChoice == 'paper') {
    if (computerChoice == 'rock') {
      humanScore++;
      return win;
    } else if (computerChoice == 'paper') {
      computerScore++;
      humanScore++;
      return tie;
    } else {
      computerScore++;
      return lose;
    }
  } else if (humanChoice == 'scissors') {
    if (computerChoice == 'rock') {
      computerScore++;
      return lose;
    } else if (computerChoice == 'paper') {
      humanScore++;
      return win;
    } else {
      computerScore++;
      humanScore++;
      return tie;
    }
  }
}

const playGame = () => {
  for(let i = 0; i < 3; i++){
    playRound(getHumanChoice(), getComputerChoice())
  }

  return humanScore > computerScore ? "You win the game!" : "You lose the game!"
}

```

