---
id: 673b567e3ba535dda140d278
title: Construir um Sistema de Votação
challengeType: 26
dashedName: build-a-voting-system
---

# --description--

Neste laboratório, você construirá um sistema de votação que usa `Map` para criar uma enquete e `Set` para evitar votos duplicados.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve inicializar uma variável `poll` para um novo objeto `Map`.

2. Você deve ter uma função `addOption` que aceita um parâmetro `option`.

3. Na função `addOption`:

   - Se a `option` ainda não existir na enquete, ela deve ser adicionada à enquete com um `Set` vazio como seu valor para rastrear os eleitores. Você também deve retornar a mensagem `Option "<option>" added to the poll.`

   - Se o `option` já existir, ele deve retornar a mensagem `Option "<option>" already exists.`.

   - Se você tentar adicionar uma opção vazia, a função deve retornar a mensagem `Option cannot be empty.`.

4. Você deve ter uma função `vote` que aceita dois parâmetros, `option` (a opção para a qual votar) e `voterId` (um ID único para o eleitor).

5. Na função `vote`:

   - Se a `option` não existir na enquete, a função deve retornar a mensagem `Option "<option>" does not exist.`.

   - Se a `option` existir, a função deve verificar se o `voterId` já votou nesta `option`.

   - Se o eleitor já votou, a função deve retornar a mensagem `Voter <voterId> has already voted for "<option>".`

   - Se o eleitor não votou, `voterId` deve ser adicionado ao `Set` de `voters` para esta opção. A função deve retornar a mensagem `Voter <voterId> voted for "<option>".`

6. Você deve ter pelo menos três opções na sua `poll`.

7. Sua `poll` deve ter pelo menos três votos.

8. Você deve ter uma função `displayResults` que retorne os resultados da pesquisa no seguinte formato:

```js
Poll Results:
OptionA: N votes
OptionB: N votes
.
.

/*
sample output

Poll Results:
Turkey: 2 votes
Morocco: 1 votes
*/
```

# --hints--

Você deve ter uma variável `poll` inicializada com um novo objeto `Map`.

```js
assert.exists(poll);
assert.instanceOf(poll, Map);
```

Você deve definir uma função `addOption` que aceita um parâmetro `option`.

```js
assert.isFunction(addOption);
assert.match(
  addOption.toString(),
  /^function\s+addOption\s*\(\s*option\s*\)\s*\{|\(?\s*option\s*\)?\s*=>/
);
```

Você deve definir uma função `vote` que aceita dois parâmetros, `option` e `voterId`.

```js
assert.isFunction(vote);
assert.match(vote.toString(), /\(\s*option\s*,\s*voterId\s*\)/);
```

Você deve definir uma função `displayResults` para exibir os resultados da enquete.

```js
assert.isFunction(displayResults);
```

Você deve ter pelo menos três opções na sua `poll`.

```js
assert.isAtLeast(poll.size, 3);
```

Sua `poll` deve ter pelo menos três votos.

```js
let totalVotes = 0;
for (let voters of poll.values()) {
  totalVotes += voters.size;
}
assert.isAtLeast(totalVotes, 3);
```

Você deve garantir que cada opção de votação corresponda a um objeto `Set`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();

  addOption('Turkey');
  assert.instanceOf(poll.get('Turkey'), Set);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`addOption("Egito")` deve retornar `Opção "Egito" adicionada à enquete.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(addOption('Egypt'), 'Option "Egypt" added to the poll.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Adicionar uma opção vazia deve retornar `"A opção não pode estar vazia."`

```js
assert.equal(addOption(''), 'Option cannot be empty.');
```

Quando `Turkey` já estiver adicionado, `addOption("Turkey")` deve retornar `Option "Turkey" already exists.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Turkey');
  assert.equal(addOption('Turkey'), 'Option "Turkey" already exists.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Quando `Malaysia` existir nas opções de votação, `vote("Malaysia", "traveler1")` deve retornar `Voter traveler1 voted for "Malaysia".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Malaysia');
  assert.equal(
    vote('Malaysia', 'traveler1'),
    'Voter traveler1 voted for "Malaysia".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`vote` deve atualizar o `Set` de eleitores para uma opção.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Alaska');
  vote('Alaska', 'seal');
  assert.isTrue(poll.get('Alaska').has('seal'));
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Quando `traveler1` tenta votar novamente em `Algeria`, `vote("Algeria", "traveler1")` deve retornar `Voter traveler1 has already voted for "Algeria".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Algeria');
  vote('Algeria', 'traveler1');
  assert.equal(
    vote('Algeria', 'traveler1'),
    'Voter traveler1 has already voted for "Algeria".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Votos duplicados não devem aumentar o tamanho do `Set`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  vote('Turkey', 'traveler1');
  const voterSet = poll.get('Turkey');
  const initialSize = voterSet.size;

  vote('Turkey', 'traveler1');
  assert.equal(voterSet.size, initialSize);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Quando `Nigeria` não estiver nas opções de votação, `vote("Nigeria", "traveler2")` deve retornar `Opção "Nigeria" não existe.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(
    vote('Nigeria', 'traveler2'),
    'Option "Nigeria" does not exist.'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Uma opção única deve ser capaz de receber múltiplos votos.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Bali');
  assert.equal(vote('Bali', 'traveler1'), 'Voter traveler1 voted for "Bali".');
  assert.equal(vote('Bali', 'traveler2'), 'Voter traveler2 voted for "Bali".');
  assert.equal(poll.get('Bali').size, 2);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`displayResults()` deve retornar os resultados no formato correto.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const poll = new Map();

function addOption(option) {
  if (!option || option.trim() === '') {
    return 'Option cannot be empty.';
  }
  if (!poll.has(option)) {
    poll.set(option, new Set());
    return `Option "${option}" added to the poll.`;
  } else {
    return `Option "${option}" already exists.`;
  }
}

function vote(option, voterId) {
  if (!poll.has(option)) {
    return `Option "${option}" does not exist.`;
  }
  const voters = poll.get(option);
  if (voters.has(voterId)) {
    return `Voter ${voterId} has already voted for "${option}".`;
  } else {
    voters.add(voterId);
    return `Voter ${voterId} voted for "${option}".`;
  }
}

function displayResults() {
  let results = 'Poll Results:\n';
  for (let [option, voters] of poll.entries()) {
    results += `${option}: ${voters.size} votes\n`;
  }
  return results.trim();
}

addOption('Turkey');
addOption('Morocco');
addOption('Spain');

vote('Turkey', 'traveler1');
vote('Turkey', 'traveler2');
vote('Morocco', 'traveler3');
```
