---
id: 657efa642593c5746acc5c81
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

`if` 文の中で、大文字の文字を小文字に変換し、この小文字の前にアンダースコアを付ける必要があります。

大文字を小文字に変換するには `.lower()` 文字列メソッドを使ってください。その後、文字の前にアンダースコアを付けます。結果は `converted_character` という名前の変数に代入してください。

# --hints--

変更した文字は `converted_character` という名前の変数に代入してください。

```js
({
    test: () => {
   assert(runPython(`
      converted_character_variable = (
          _Node(_code)
          .find_function("convert_to_snake_case")
          .find_for_loops()[0]
          .find_ifs()[0]
          .find_variable("converted_character")
      )
      
      converted_character_variable.is_equivalent(
          "converted_character = '_' + char.lower()"
      ) or converted_character_variable.is_equivalent(
          "converted_character = f'_{char.lower()}'"
      )
    `)); 
    }
})
```

`pass` 文の中に `if` を含めてはいけません。

```js
({
    test: () => {
        assert(runPython(`not _Node(_code).find_function("convert_to_snake_case").find_for_loops()[0].find_ifs()[0].has_stmt("pass")`)); 
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            pass
--fcc-editable-region--
```
