---
id: 657e928716b77b2277980276
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このプロジェクトでは、`camelCase` または `PascalCase` 形式の文字列を入力として受け取り、2つの方法で `snake_case` 形式の文字列に変換するプログラムを作成します。まずは for ループを使い、その後リスト内包表記を使って同じ結果を得ます。リスト内包表記がコードをどれだけ簡潔にできるかがわかります。

`convert_to_snake_case()` という名前の文字列を入力として受け取る `pascal_or_camel_cased_string` という新しい関数を定義し始めてください。今のところ、関数内に `pass` 文を追加してください。

# --hints--

`convert_to_snake_case()` という名前の新しい関数を定義してください。最後のコロンを忘れずに、関数の本体は `pass` で埋めてください。

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(convert_to_snake_case)
    `));
  }
})
```

関数は `pascal_or_camel_cased_string` という名前のパラメータを受け取る必要があります。

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      'pascal_or_camel_cased_string' in inspect.signature(convert_to_snake_case).parameters
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
