---
id: 657ef2a86d4e545cec9a85fb
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

キャメルケースとパスカルケースの両方で、大文字は新しい単語の始まりを示します。入力された文字列をスネークケースに変換するには、入力文字列の文字が大文字かどうかを確認する必要があります。

文字が大文字かどうかを確認するには、`.isupper()` 文字列のメソッドを使えます。このメソッドは、文字が大文字の場合は `True` を返し、大文字でない場合は `False` を返します。

`for` ループの中で、現在の文字が大文字かどうかを確認する `if` 文を追加してください。`pass` 文は新しい `if` 文の中に移動させてください。

# --hints--

条件に char.isupper() を使った新しい if文 を書く必要があります。最後にコロンを付け、if文の本体は pass キーワードで埋めてください。

```js
({
    test: () => {        
        const commentless_code = __helpers.python.removeComments(code)
        const transformedCode = commentless_code.replace(/\r/g, "");
        const {function_body} = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        assert.match(function_body, /\s+if\s+char\.isupper\s*\(\s*\)\s*:\s+pass/);
    }
})
```

ループの本体にある `pass` 文を `if` 文に置き換えてください。

```js
({
   test: () => {
        const commentless_code = __helpers.python.removeComments(code)
        const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+pascal_or_camel_cased_string\s*/);
        assert.notMatch(block_body, /^\s+pass/);

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        pass
--fcc-editable-region--        
```
