---
id: 657ed53c19461d4b95c4757a
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

文字列をスネークケースに変換した後の文字を格納する空のリストを追加する必要があります。

関数内で、`pass` の文を空のリスト `snake_cased_char_list` を作成するコードに置き換えてください。

# --hints--

`pass` の中で `snake_cased_char_list` を空のリスト `convert_to_snake_case` に置き換える必要があります。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const commentless_code = __helpers.python.removeComments(transformedCode)
        const convert_to_snake_case = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;
        assert.notMatch(function_body, /\bpass\b/);
        assert.match(function_body, / +snake_cased_char_list\s*=\s*\[\s*\]\s*/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    pass    
--fcc-editable-region--
```
