---
id: 657ee28cefc4945568287673
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

空のリストが用意できたので、次は入力文字列を繰り返し処理してスネークケースに変換していきましょう。

関数内で、先ほど作成したリストの下に、`for` ループを追加して `pascal_or_camel_cased_string` を繰り返し処理してください。ターゲット変数の名前は `char` にしてください。今はループ本体に `pass` 文をプレースホルダーとして追加しましょう。

# --hints--

新しい `for` ループを書き、ターゲット変数の名前は `char` にしてください。最後のコロンを忘れずに、ループ本体には `pass` を使ってください。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +for\s+char\s+in\s+pascal_or_camel_cased_string\s*:\s*pass[\s]*$/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
--fcc-editable-region--
```
