---
id: 657efce98e958b75df86b305
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

`if` 文の中で、変換した文字を先に作成したリストに追加します。

これには `.append()` メソッドを使います。このメソッドは、呼び出されたリストの末尾に指定したオブジェクトを追加します。

`.append()` に対して `snake_cased_char_list` メソッドを使い、`converted_character` をリストに追加してください。

# --hints--

変換した文字を `.append()` に追加するには `snake_cased_char_list` メソッドを使うべきです。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_char_list\.append\s*\(\s*converted_character\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            converted_character = '_' + char.lower()
--fcc-editable-region--
```
