---
id: 657effaa2a5e0277d71f9cbe
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

この時点で、変数 `snake_cased_char_list` は変換された文字のリストを保持しています。これらの文字を1つの文字列に結合するには、`.join()` メソッドを利用できます。

`join` メソッドは、リストの各要素を指定した文字列（セパレーター）で区切って連結し、1つの文字列にします。

```py
result_string = ''.join(characters)
```

上の例では、`characters` リストの要素を空文字列をセパレーターとして連結し、1つの文字列にまとめています。

さて、`for` ループの直後に、`.join()` メソッドを使って `snake_cased_char_list` の要素を空文字列で結合してください。結果は `snake_cased_string` という新しい変数に代入します。


# --hints--

`snake_cased_char_list` リストの文字を `.join()` メソッドで文字列に結合し、`snake_cased_string` ループの後に `for` という変数に代入してください。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_string\s*=\s*('|")\1\.join\s*\(\s*snake_cased_char_list\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--


--fcc-editable-region--
```
