---
id: 657f456223b8c1187b461987
title: ステップ 19
challengeType: 20
dashedName: step-19
---

# --description--

Pythonでは、<dfn>リスト内包表記</dfn>は、既存のイテラブルの各要素に式を適用し、条件で要素をフィルターしながら新しいリストを生成できる構文です。短く書けるだけでなく、リスト内包表記はしばしば高速に動作します。

基本的なリスト内包表記は、式の後に`for`節が続きます。

```py
spam = [i * 2 for i in iterable]
```

上記では、変数`i`を使って`iterable`を繰り返しています。結果のリストの各要素は、現在の反復で式`i * 2`を評価して得られます。

このステップでは、空のリスト`snake_cased_char_list`をリスト内包表記の構文を使って埋める必要があります。

空のリストを、`pascal_or_camel_cased_string`の各文字を小文字に変換し、先頭にアンダースコアを付けるリスト内包表記に変えてください（以前コメントアウトしたコードが式を書くのに役立つかもしれません）。`char`を使って`pascal_or_camel_cased_string`を繰り返してください。

# --hints--

`snake_cased_char_list`を`pascal_or_camel_cased_string`を繰り返すリスト内包表記に変える必要があります。

```js
({ 
    test: () => assert(runPython(`
    iters = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_iters()
    len(iters) == 1 and iters[0].is_equivalent("pascal_or_camel_cased_string")
    `)) 
})
```

リスト内包表記では`char`を使って`pascal_or_camel_cased_string`を繰り返してください。

```js
({ 
    test: () => assert(runPython(`
    targets = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_targets()
    len(targets) == 1 and targets[0].is_equivalent("char")
    `)) 
})
```

リスト内包表記は、`'_' + char.lower()`の各`char`に対して`pascal_or_camel_cased_string`を評価する必要があります。

```js
({ 
    test: () => assert(runPython(`
    target_var = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list")
    target_var.find_comp_expr().is_equivalent("'_' + char.lower()") and len(target_var.find_comp_ifs()) == 0
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
