---
id: 657f425dbab54e11993c80f0
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

これまでのプロジェクトでは、`for` ループを使って入力文字列を繰り返し処理し、目的の出力に変換してきました。ここからは `for` ループからリスト内包表記への移行を始めます。

まず、`convert_to_snake_case()` 関数内のすべてのコード行をコメントアウトしてください。削除しないでください。リスト内包表記でロジックを実装するときに役立ちます。

テスト中にコードが失敗しないように、関数本体に `pass` キーワードを追加することを忘れないでください。

# --hints--

`convert_to_snake_case()` 関数内の10行のコードすべてをコメントアウトし、それ以外は変更しないでください。エラーを避けるために、関数本体に一時的に `pass` を追加してください。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "").replace(/\n#/g, "\n #");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.equal(function_body.split("#").length - 1, 10);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()

```
