---
id: 657f04ed0035f47ed04d0f1f
title: ステップ 13
challengeType: 20
dashedName: step-13
---

# --description--

`convert_to_snake_case()` 関数の出力を表示するには、`main()` 関数を呼び出す必要があります。

既存の2つの関数と同じレベルに、`main()` 関数の呼び出しを追加してください。実行すると、与えられたキャメルケースまたはパスカルケースの文字列がスネークケースに変換されているのが確認できます。

# --hints--

`main()` 関数の呼び出しを追加してください。

```js
({
  test: () => {
    const pythonCode = `_Node(_code.replace('\\\\r', '')).has_call("main()")`;
    const result = runPython(pythonCode);
    assert(result);
  }
});
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

--fcc-editable-region--
```
