---
id: 657f47b12c51e41b3149e584
title: ステップ 21
challengeType: 20
dashedName: step-21
---

# --description--

それでも、最終的な結果は完全に望んだものではありません。`if` 節で除外された文字に対しては、別の式を実行する必要があります。そのために `else` 節を使います。

```py
spam = [i * 2 if i > 0 else -1 for i in iterable]
```

`if` 節とは異なり、`if`/`else` 構文は式と `for` キーワードの間に置く必要があることに注意してください。

大文字でない文字は変更せずにそのままにするように、リスト内包表記を修正してください。

# --hints--

リスト内包表記を修正して、`'_' + char.lower()` の場合は式 `char.isupper()` を評価し、それ以外は char を返すようにしてください。

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_expr().is_equivalent("'_' + char.lower() if char.isupper() else char")
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string if char.isupper()]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
