---
id: 657f4add33ea4b1f61ba3dc8
title: ステップ 23
challengeType: 20
dashedName: step-23
---

# --description--

最後に、この新しい実装をプログラムを実行して試してみましょう。入力文字列を`'IAmAPascalCasedString'`に変更し、それが`'i_am_a_pascal_cased_string'`として出力されるか確認してください。たとえそれが嘘でもです。

すべて正しくできていれば、入力文字列が以前と同じようにスネークケースに変換されているはずです。

おめでとうございます！これで`convert_to_snake_case()`関数の準備が整いました。

# --hints--

`'aLongAndComplexString'`関数内で入力文字列を`'IAmAPascalCasedString'`から`main()`に変更してください。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\(\s*convert_to_snake_case\(\s*'IAmAPascalCasedString'\s*\)\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

    
--fcc-editable-region--

main()
```

# --solutions--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('IAmAPascalCasedString'))

main()
```
