---
id: 657f025ec86c3d7c4177b6be
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

関数が完成したので、別の関数の中でそれを使うことができます。

`main()` という新しい関数を作成し、関数の本体として `pass` を使ってください。

# --hints--

`main()` という名前の新しい関数を定義してください。最後のコロンを忘れずに、関数の本体には `pass` を使いましょう。

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(main)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--


--fcc-editable-region--
```
