---
id: 657f0044be09db790b1eb1c5
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

パスカルケースでは、文字列は大文字で始まります。すべての文字を小文字に変換し、アンダースコアを追加すると、文字列の先頭に余分なアンダースコアが付く場合があります。

これを修正する最も簡単な方法は、`.strip()` 文字列メソッドを使うことです。このメソッドは、引数として渡された文字の集合の中から、文字列の先頭または末尾にある文字を削除します。例えば：

```py
original_string = "_example_string_"

clean_string = original_string.strip('_')
```

`strip()` メソッドは `original_string` に適用されます。これにより、先頭と末尾のアンダースコアが削除されます。上記の例の結果は文字列 `'example_string'` になります。

`clean_snake_cased_string` という名前の新しい変数を宣言し、`.strip()` に `snake_cased_string` メソッドを適用した結果を代入してください。このとき、メソッドの引数として `'_'` を渡します。

# --hints--

`snake_cased_string.strip('_')` の結果の文字列を `clean_snake_cased_string` という名前の変数に代入してください。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +clean_snake_cased_string\s*=\s*snake_cased_string.strip\(\s*("|')_\1\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--
    snake_cased_string = ''.join(snake_cased_char_list)
--fcc-editable-region--
```
