---
id: 657f465f8e718b19c5105ae5
title: ステップ 20
challengeType: 20
dashedName: step-20
---

# --description--

リスト内包表記は条件文を受け入れ、特定の条件が満たされた場合にのみ指定された式を評価します。

```py
spam = [i * 2 for i in iterable if i > 0]
```

出力からわかるように、`pascal_or_camel_cased_string` から生成された文字のリストが結合されています。リスト内包表記の中の式は各文字に対して評価されるため、結果はすべての文字がアンダースコアで区切られた小文字の文字列になります。

上記の例に従って、リスト内包表記に `if` 節を追加し、文字が大文字の場合にのみ式が実行されるようにしてください。

# --hints--

リスト内包表記に `if` 節を、条件 `char.isupper()` で追加してください。

```js
({ 
    test: () => assert(runPython(`
    ifs = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_ifs()
    len(ifs) == 1 and ifs[0].is_equivalent("char.isupper()")
    `))
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
