---
id: 663b16e62fee463b4caf46e9
title: ステップ 18
challengeType: 20
dashedName: step-18
---

# --description--

リスト `snake_cased_char_list` の要素を結合した後、結果の文字列から先頭および末尾のアンダースコアを取り除く必要があります。そのために、アンダースコア文字 `strip` を引数にして `_` メソッドを使ってください。

メソッドの呼び出しはチェーンできます。つまり、あるメソッド呼び出しの結果を別のメソッド呼び出しの対象オブジェクトとして使えます。

```py
words_list = ['hello', 'world', 'this', 'is', 'chained', 'methods']
result = ' '.join(words_list).upper()
```

上の例では、`.upper()` メソッドが `' '.join(words_list)` にチェーンされているため、`.upper()` は `.join()` 呼び出しの結果に対して呼び出されます。

`return` 文を修正して、先頭および末尾のアンダースコアを取り除くために `''.join(snake_cased_char_list)` メソッドの呼び出しを `.strip()` にチェーンしてください。

# --hints--

return文を修正して、アンダースコア文字 `strip` を引数にした `_` メソッドを含めてください。

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, /return\s*''.join\(snake_cased_char_list\).strip\('_'\)/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list)
--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
