---
id: 663b10c10a4c0a0e095137ee
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

リスト `snake_cased_char_list` を返す代わりに、その要素を空の文字列 `''` を区切り文字として結合した単一の文字列にする必要があります。

`return` の文を修正して、`snake_cased_char_list` を空の文字列を区切り文字として結合した結果を返してください。

# --hints--

`return` の文を修正して、`snake_cased_char_list` を空の文字列 `''` を区切り文字として `join()` メソッドで結合した結果を返すようにしてください。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_function('convert_to_snake_case').find_return().is_equivalent('return "".join(snake_cased_char_list)')
    `))
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--
    return snake_cased_char_list
--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
