---
id: 66325a250690a3612c1db0f6
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

CamperBot は `bgHexCodeSpanElement` 要素の `span` という `id` への参照を保存するために `bg-hex-code` という新しい変数を作成しました。しかし、その変数をコンソールにログしようとすると、`null` が表示されます。

`null` は JavaScript における特別な値で、値が存在しないことを表します。これは存在しないオブジェクトのプロパティにアクセスしようとした場合に起こります。

この場合、CamperBot は `document.querySelector` メソッドに正しいセレクターを渡していません。

`document.querySelector("bg-hex-code")` の行を修正して、`id` という `bg-hex-code` を持つ要素を正しく選択するようにしてください。

# --hints--

`"bg-hex-code"` を修正して、`id` の `bg-hex-code` を正しく参照するようにしてください。CSS の `id` セレクターはハッシュ記号（`#`）で始まることを覚えておいてください。

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Build a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Random Background Color changer</h1>

    <main>
      <section class="bg-information-container">
        <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
      </section>

      <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --yellow: #f1be32;
  --golden-yellow: #feac32;
  --dark-purple: #110815;
  --light-grey: #efefef;
}

body {
  background-color: var(--dark-purple);
  color: var(--light-grey);
  text-align: center;
}

.bg-information-container {
  margin: 15px 0 25px;
  font-size: 1.2rem;
}

.btn {
  cursor: pointer;
  padding: 10px;
  margin: 10px;
  color: var(--dark-purple);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];
function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

--fcc-editable-region--
const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("bg-hex-code");

console.log(bgHexCodeSpanElement);
--fcc-editable-region--
```
