---
id: 6650c9a94d6e13d14a043a69
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

CamperBot はランダムな背景色を変更する機能を作ろうとしています。しかし、問題が発生していて、コードのデバッグを手伝ってほしいようです。

CamperBot はすでにプロジェクトの HTML と CSS を追加していますが、なぜかページにスタイルやコンテンツが表示されないことに困っています。

コンソールを開くと、次のメッセージが表示されます：

```js
SyntaxError: unknown: Unexpected token, expected "," (5:2)
```

JavaScript エンジンが解釈できないものに遭遇すると構文エラーがスローされます。この場合、CamperBot の `darkColorsArr` 配列に構文エラーがあるようです。

`darkColorsArr` 配列の構文エラーを修正すれば、ページにコンテンツとスタイルが表示されるはずです。

# --hints--

`darkColorsArr` 配列の構文エラーを解決してください。配列の要素はカンマで区切る必要があることを覚えておいてください。

```js
assert.isArray(darkColorsArr);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Build a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Random Background Color changer</h1>

    <main>
      <section class="bg-information-container">
        <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
      </section>

      <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --yellow: #f1be32;
  --golden-yellow: #feac32;
  --dark-purple: #110815;
  --light-grey: #efefef;
}

body {
  background-color: var(--dark-purple);
  color: var(--light-grey);
  text-align: center;
}

.bg-information-container {
  margin: 15px 0 25px;
  font-size: 1.2rem;
}

.btn {
  cursor: pointer;
  padding: 10px;
  margin: 10px;
  color: var(--dark-purple);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

```

```js
--fcc-editable-region--
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C"
  "#616A6B"
  "#4A235A"
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];
--fcc-editable-region--
```
