---
id: 66fbcf750a62784cf11f5632
title: الخطوة 8
challengeType: 1
dashedName: step-8
---

# --description--

أنشئ دالة `getTotalIngredients` تأخذ معلمة واحدة تمثل مصفوفة تحتوي على المكونات، وتُرجع عدد المكونات في المصفوفة التي تم تمريرها إلى الدالة.

# --hints--

يجب أن تنشئ دالة `getTotalIngredients`.

```js
assert.isFunction(getTotalIngredients);
```

يجب أن تحتوي دالة `getTotalIngredients` على معلمة واحدة.

```js
assert.lengthOf(getTotalIngredients, 1);
```

يجب أن تُرجع دالة `getTotalIngredients` رقمًا.

```js
assert.isNumber(getTotalIngredients(recipe1.ingredients));
```

يجب أن تُرجع `getTotalIngredients(recipe1.ingredients)` القيمة `4`.

```js
assert.strictEqual(getTotalIngredients(recipe1.ingredients), 4);
```

يجب أن تُرجع `getTotalIngredients(recipe2.ingredients)` القيمة `5`.

```js
assert.strictEqual(getTotalIngredients(recipe2.ingredients), 5);
```

يجب أن تُرجع `getTotalIngredients(recipe3.ingredients)` القيمة `3`.

```js
assert.strictEqual(getTotalIngredients(recipe3.ingredients), 3);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

--fcc-editable-region--

--fcc-editable-region--
```
