---
id: 66fbcf750a62784cf11f5633
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

أنشئ دالة `getDifficultyLevel` تأخذ رقمًا يشير إلى وقت الطهي كمعلمة.

إذا كان وقت الطهي أقل من أو يساوي `30`، يجب أن تُرجع الدالة القيمة `"easy"`. إذا كان أقل من أو يساوي `60`، يجب أن تُرجع الدالة القيمة `"medium"`. وإلا، يجب أن تُرجع الدالة القيمة `"hard"`.

# --hints--

يجب أن تنشئ دالة `getDifficultyLevel`.

```js
assert.isFunction(getDifficultyLevel);
```

يجب أن تحتوي دالة `getDifficultyLevel` على معلمة واحدة فقط.

```js
assert.lengthOf(getDifficultyLevel, 1);
```

يجب أن تُرجع دالة `getDifficultyLevel` القيمة `"easy"` عندما يكون وقت الطهي أقل من أو يساوي `30`.

```js
assert.strictEqual(getDifficultyLevel(10), "easy");
assert.strictEqual(getDifficultyLevel(20), "easy");
assert.strictEqual(getDifficultyLevel(29), "easy");
assert.strictEqual(getDifficultyLevel(30), "easy");
```

يجب أن تُرجع دالة `getDifficultyLevel` القيمة `"medium"` عندما يكون وقت الطهي أكبر من `30` وأقل من أو يساوي `60`.

```js
assert.strictEqual(getDifficultyLevel(31), "medium");
assert.strictEqual(getDifficultyLevel(40), "medium");
assert.strictEqual(getDifficultyLevel(50), "medium");
assert.strictEqual(getDifficultyLevel(60), "medium");
```

يجب أن تُرجع دالة `getDifficultyLevel` القيمة `"hard"` عندما يكون وقت الطهي أكبر من `60`.

```js
assert.strictEqual(getDifficultyLevel(61), "hard");
assert.strictEqual(getDifficultyLevel(65), "hard");
assert.strictEqual(getDifficultyLevel(70), "hard");
assert.strictEqual(getDifficultyLevel(75), "hard");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

--fcc-editable-region--

--fcc-editable-region--
```
