---
id: 66fbcf750a62784cf11f5634
title: الخطوة 10
challengeType: 1
dashedName: step-10
---

# --description--

حان الوقت لاختبار كل دالة. يمكنك استخدام أي وصفة لهذا، لكن هذا الدرس سيبدأ بـ `recipe1`.

أنشئ متغيرين جديدين: `recipe1TotalIngredients` و `recipe1DifficultyLevel`. عيّن لهما القيم باستدعاء الدالة المناسبة لكل متغير وتمرير الخاصية المناسبة من `recipe1`.

أخيرًا، سجّل كل متغير في وحدة التحكم لرؤية النتائج.

# --hints--

يجب أن تنشئ متغيرًا باسم `recipe1TotalIngredients`.

```js
assert.isNotNull(recipe1TotalIngredients);
```

يجب تعيين متغير `recipe1TotalIngredients` إلى `getTotalIngredients(recipe1.ingredients)`.

```js
assert.equal(recipe1TotalIngredients, getTotalIngredients(recipe1.ingredients));
```

يجب أن تسجل `recipe1TotalIngredients` في وحدة التحكم.

```js
assert.match(code, /console\.log\(recipe1TotalIngredients\);?/);
```

يجب أن تنشئ متغيرًا باسم `recipe1DifficultyLevel`.

```js
assert.isNotNull(recipe1DifficultyLevel);
```

يجب تعيين متغير `recipe1DifficultyLevel` إلى `getDifficultyLevel(recipe1.cookingTime)`.

```js
assert.equal(recipe1DifficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

يجب أن تسجل `recipe1DifficultyLevel` في وحدة التحكم.

```js
assert.match(code, /console\.log\(recipe1DifficultyLevel\);?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
