---
id: 66fbcf750a62784cf11f562d
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

أنشئ كائنًا باسم `recipe1`. داخل كائن `recipe1`، أنشئ خاصية `name` بقيمة `"Spaghetti Carbonara"`.

أيضًا داخل كائن `recipe1`، أنشئ خاصية `ingredients` بقيمة مصفوفة. يجب أن تحتوي المصفوفة على `spaghetti` و`Parmesan cheese` و`pancetta` و`black pepper`.

# --hints--

يجب أن تنشئ كائنًا باسم `recipe1`.

```js
assert.isObject(recipe1);
```

يجب أن يحتوي كائن `recipe1` على خاصية `name`.

```js
assert.property(recipe1, 'name');
```

يجب تعيين خاصية `name` إلى `Spaghetti Carbonara`.

```js
assert.equal(recipe1.name, 'Spaghetti Carbonara');
```

يجب أن يحتوي كائن `recipe1` على خاصية `ingredients`.

```js
assert.property(recipe1, 'ingredients');
```

يجب أن تكون قيمة خاصية `ingredients` مصفوفة.

```js
assert.isArray(recipe1.ingredients);
```

يجب أن تحتوي قيمة المصفوفة في خاصية `ingredients` على `spaghetti` و`Parmesan cheese` و`pancetta` و`black pepper`.

```js
assert.deepEqual(recipe1.ingredients, ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

--fcc-editable-region--

--fcc-editable-region--
```
