---
id: 66fbcdfed6848e48947dbbc9
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه ورشة العمل، ستنشئ متتبع وصفات باستخدام كائنات JavaScript.

ابدأ بإنشاء مصفوفة فارغة باسم `recipes`. هذه هي المصفوفة التي ستُدفع إليها كائنات الوصفات لاحقًا.

# --hints--

يجب إنشاء مصفوفة باسم `recipes`.

```js
assert.isArray(recipes);
```

يجب أن تكون مصفوفة `recipes` فارغة.

```js
assert.isEmpty(recipes);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
