---
id: 66fbcf750a62784cf11f562e
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

أنشئ كائنًا باسم `recipe2` يحتوي على الخصائص والقيم التالية:

| المفتاح | القيمة   |
| ----------- | ------- |
| `name`       | `Chicken Curry` |
| `ingredients`       | `["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]` |
| `cookingTime`       | `42` |
| `totalIngredients`   | `null`   |
| `difficultyLevel`   | `""`   |

# --hints--

يجب أن تنشئ كائنًا اسمه `recipe2`.

```js
assert.isObject(recipe2);
```

يجب أن يحتوي كائن `recipe2` على خاصية `name`.

```js
assert.property(recipe2, "name");
```

يجب تعيين خاصية `name` إلى `Chicken Curry`.

```js
assert.equal(recipe2.name, "Chicken Curry");
```

يجب أن يحتوي كائن `recipe2` على خاصية `ingredients`.

```js
assert.property(recipe2, "ingredients");
```

يجب أن تكون قيمة خاصية `ingredients` مصفوفة.

```js
assert.isArray(recipe2.ingredients);
```

يجب أن تحتوي المصفوفة في خاصية `ingredients` على `chicken breast` و`coconut milk` و`curry powder` و`onion` و`garlic`.

```js
assert.deepEqual(recipe2.ingredients, ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]);
```

يجب أن يحتوي كائن `recipe2` على خاصية `cookingTime`.

```js
assert.property(recipe2, "cookingTime");
```

يجب أن تكون قيمة خاصية `cookingTime` رقمًا.

```js
assert.isNumber(recipe2.cookingTime);
```

يجب تعيين خاصية `cookingTime` إلى `42`.

```js
assert.equal(recipe2.cookingTime, 42);
```

يجب أن يحتوي كائن `recipe2` على خاصية `totalIngredients`.

```js
assert.property(recipe2, "totalIngredients");
```

يجب تعيين خاصية `totalIngredients` إلى `null`.

```js
assert.isNull(recipe2.totalIngredients);
```

يجب أن يحتوي كائن `recipe2` على خاصية `difficultyLevel`.

```js
assert.property(recipe2, "difficultyLevel");
```

يجب تعيين خاصية `difficultyLevel` إلى سلسلة نصية فارغة.

```js
assert.deepEqual(recipe2.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
