---
id: 66fbcf750a62784cf11f5630
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

يجب عليك الآن دفع الكائنات الثلاثة إلى مصفوفة `recipes`. للقيام بذلك، يمكنك استخدام الطريقة `push()`.

استخدم الطريقة `push()` لدفع جميع كائنات الوصفات إلى مصفوفة `recipes`. تأكد من دفع `recipe1`، و`recipe2`، و`recipe3` بهذا الترتيب.

احذف أيضًا المتغيرات `recipe1Name`، و`recipe2CookingTime`، و`recipe3Ingredients`، وعبارات `console.log` التي تسجل تلك المتغيرات.

# --hints--

يجب دفع جميع كائنات الوصفات إلى مصفوفة `recipes`. تأكد من دفعها بالترتيب الذي تم الإعلان عنه.

```js
assert.deepEqual(recipes, [recipe1, recipe2, recipe3]);
```

يجب إزالة المتغير `recipe1Name`.

```js
assert.throws(() => recipe1Name, ReferenceError);
```

يجب إزالة العبارة `console.log(recipe1Name);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1Name\s*\)\s*;?/);
```

يجب إزالة المتغير `recipe2CookingTime`.

```js
assert.throws(() => recipe2CookingTime, ReferenceError);
```

يجب إزالة العبارة `console.log(recipe2CookingTime);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2CookingTime\s*\)\s*;?/);
```

يجب إزالة المتغير `recipe3Ingredients`.

```js
assert.throws(() => recipe3Ingredients, ReferenceError);
```

يجب إزالة العبارة `console.log(recipe3Ingredients);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe3Ingredients\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--
const recipe1Name = recipe1.name;
console.log(recipe1Name);

const recipe2CookingTime = recipe2.cookingTime;
console.log(recipe2CookingTime);

const recipe3Ingredients = recipe3.ingredients;
console.log(recipe3Ingredients);
--fcc-editable-region--
```
