---
id: 66ff02792b3042a2b811f475
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

أضف الخصائص التالية إلى كائن `recipe1`:

| المفتاح | القيمة   |
| ----------- | ------- |
|`cookingTime`|`22`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

الخصائص التي قيمتها `null` والسلسلة النصية الفارغة ستُحدَّث لاحقًا بعد حسابها.

# --hints--

يجب أن يحتوي كائن `recipe1` على خاصية `cookingTime`.

```js
assert.property(recipe1, "cookingTime");
```

يجب تعيين خاصية `cookingTime` إلى القيمة `22`.

```js
assert.equal(recipe1.cookingTime, 22);
```

يجب أن يحتوي كائن `recipe1` على خاصية `totalIngredients`.

```js
assert.property(recipe1, "totalIngredients");
```

يجب تعيين خاصية `totalIngredients` إلى القيمة `null`.

```js
assert.isNull(recipe1.totalIngredients);
```

يجب أن يحتوي كائن `recipe1` على خاصية `difficultyLevel`.

```js
assert.property(recipe1, "difficultyLevel");
```

يجب تعيين خاصية `difficultyLevel` إلى سلسلة نصية فارغة.

```js
assert.deepEqual(recipe1.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
--fcc-editable-region--
  
--fcc-editable-region--
};
```
