---
id: 670e4f45f7116c0f216a5177
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

قبل المتابعة، يجب أن تتدرب على كيفية الوصول إلى الخصائص من كائن.

يمكنك استخدام إما النقطة (`.`) أو الأقواس (`[]`) للقيام بذلك. إليك مثالاً:

```js
const person = {
  name: "John",
  age: 30,
  job: "Software Engineer"
};

console.log(person.name); // John
console.log(person['age']);  // 30
```

احصل على خاصية `name` من `recipe1`، وعيّنها إلى المتغير `recipe1Name`.

بعد ذلك، احصل على خاصية `cookingTime` من `recipe2` وعيّنها إلى المتغير `recipe2CookingTime`.

أخيرًا، احصل على خاصية `ingredients` من `recipe3` وعيّنها إلى المتغير `recipe3Ingredients`.

تأكد من تسجيل جميع المتغيرات التي أنشأتها في وحدة التحكم.

# --hints--

يجب أن تنشئ متغيرًا باسم `recipe1Name`.

```js
assert.isDefined(recipe1Name);
```

يجب تعيين قيمة الخاصية `name` من `recipe1` إلى المتغير `recipe1Name`.

```js
assert.strictEqual(recipe1Name, recipe1.name);
```

يجب أن تسجل `recipe1Name` في وحدة التحكم.

```js
assert.match(code, /console\.log\(\s*recipe1Name\s*\)/);
```

يجب أن تنشئ متغيرًا باسم `recipe2CookingTime`.

```js
assert.isDefined(recipe2CookingTime);
```

يجب أن تعيّن قيمة خاصية `cookingTime` من `recipe2` إلى متغير `recipe2CookingTime`.

```js
assert.strictEqual(recipe2CookingTime, recipe2.cookingTime);
```

يجب أن تسجل `recipe2CookingTime` في وحدة التحكم.

```js
assert.match(code, /console\.log\(recipe2CookingTime\)/);
```

يجب أن تنشئ متغيرًا باسم `recipe3Ingredients`.

```js
assert.isDefined(recipe3Ingredients);
```

يجب أن تعيّن قيمة خاصية `ingredients` من `recipe3` إلى متغير `recipe3Ingredients`.

```js
assert.strictEqual(recipe3Ingredients, recipe3.ingredients);
```

يجب أن تسجل `recipe3Ingredients` في وحدة التحكم.

```js
assert.match(code, /console\.log\(recipe3Ingredients\)/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
