---
id: 69161f4c0c9c8c48b3c843c5
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

أنشئ كائنًا باسم `recipe3` يحتوي على الخصائص والقيم التالية:

|المفتاح|القيمة|
|-----------|-------|
|`name`|`Vegetable Stir Fry`|
|`ingredients`|`["broccoli", "carrot", "bell pepper"]`|
|`cookingTime`|`15`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

# --hints--

يجب أن تنشئ كائنًا اسمه `recipe3`.

```js
assert.isObject(recipe3);
```

يجب أن يحتوي كائن `recipe3` على خاصية `name`.

```js
assert.property(recipe3, "name");
```

يجب تعيين خاصية `name` إلى `Vegetable Stir Fry`.

```js
assert.equal(recipe3.name, "Vegetable Stir Fry");
```

يجب أن يحتوي كائن `recipe3` على خاصية `ingredients`.

```js
assert.property(recipe3, "ingredients");
```

يجب أن تكون قيمة خاصية `ingredients` مصفوفة.

```js
assert.isArray(recipe3.ingredients);
```

يجب أن تحتوي المصفوفة في خاصية `ingredients` على `broccoli` و`carrot` و`bell pepper`.

```js
assert.deepEqual(recipe3.ingredients, ["broccoli", "carrot", "bell pepper"]);
```

يجب أن يحتوي كائن `recipe3` على خاصية `cookingTime`.

```js
assert.property(recipe3, "cookingTime");
```

يجب أن تكون قيمة خاصية `cookingTime` رقمًا.

```js
assert.isNumber(recipe3.cookingTime);
```

يجب تعيين خاصية `cookingTime` إلى `15`.

```js
assert.equal(recipe3.cookingTime, 15);
```

يجب أن يحتوي كائن `recipe3` على خاصية `totalIngredients`.

```js
assert.property(recipe3, "totalIngredients");
```

يجب تعيين خاصية `totalIngredients` إلى `null`.

```js
assert.isNull(recipe3.totalIngredients);
```

يجب أن يحتوي كائن `recipe3` على خاصية `difficultyLevel`.

```js
assert.property(recipe3, "difficultyLevel");
```

يجب تعيين خاصية `difficultyLevel` إلى سلسلة نصية فارغة.

```js
assert.deepEqual(recipe3.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
