---
id: 69aaafeaa4037197c75587cd
title: 實作一個 `Queue`
challengeType: 26
dashedName: implement-a-queue
---

# --description--

在本實驗中，你將使用函式實作一個佇列資料結構。佇列是一種先進先出（FIFO）的資料結構，元素會被添加到佇列的尾端，並從前端移除。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `enqueue` 函式（程式），用來將元素添加到佇列的尾端。
1. 你應該有一個 `dequeue` 函式，能移除並傳回佇列的前端元素。
1. 你應該有一個 `front` 函式，該函式會傳回佇列的前端元素而不將其移除。
1. 你應該有一個 `size` 函式，傳回佇列中的元素數量。
1. 你應該有一個 `isEmpty` 函式（程式），如果佇列為空則傳回 `true`，否則傳回 `false`。

# --hints--

你應該有一個 `enqueue` 函式（程式）。

```js
assert.isFunction(enqueue);
```

`enqueue` 函式應該將元素添加到佇列的尾端。

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

你應該有一個 `dequeue` 函式（程式）。

```js
assert.isFunction(dequeue);
```

`dequeue` 函式應該移除並傳回佇列的前端元素。

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(dequeue(queue), 'first');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(size(queue), 1);
```

當在空的佇列上 `dequeue` 時，應該傳回 `undefined`。

```js
const queue = initQueue();
assert.isUndefined(dequeue(queue));
```

`enqueue` 應該在 `dequeues` 之後繼續將元素添加到佇列的尾端。

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(dequeue(queue), 'first');
enqueue(queue, 'third');
assert.strictEqual(front(queue), 'second');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(dequeue(queue), 'third');
```

不同的佇列實體不應該共用狀態。

```js
const queueOne = initQueue();
const queueTwo = initQueue();
enqueue(queueOne, 'first');
enqueue(queueTwo, 'second');
assert.strictEqual(front(queueOne), 'first');
assert.strictEqual(front(queueTwo), 'second');
assert.strictEqual(size(queueOne), 1);
assert.strictEqual(size(queueTwo), 1);
```

你應該有一個 `front` 函式（程式）。

```js
assert.isFunction(front);
```

`front` 函式應該傳回佇列的前端元素而不將其移除。

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

當在空的佇列上呼叫 `front` 時，應該傳回 `undefined`。

```js
const queue = initQueue();
assert.isUndefined(front(queue));
```

你應該有一個 `size` 函式（程式）。

```js
assert.isFunction(size);
```

`size` 函式應該傳回佇列中正確的元素數量。

```js
const queue = initQueue();
assert.strictEqual(size(queue), 0);
enqueue(queue, 'first');
assert.strictEqual(size(queue), 1);
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(size(queue), 3);
dequeue(queue);
assert.strictEqual(size(queue), 2);
```

你應該有一個 `isEmpty` 函式（程式）。

```js
assert.isFunction(isEmpty);
```

`isEmpty` 函式應該為空佇列傳回 `true`。

```js
const queue = initQueue();
assert.isTrue(isEmpty(queue));
```

`isEmpty` 函式（程式）對於非空佇列應該傳回 `false`。

```js
const queue = initQueue();
enqueue(queue, 'element');
assert.isFalse(isEmpty(queue));
```

`isEmpty` 函式（程式）應該在所有元素從佇列中移除後傳回 `true`。

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
dequeue(queue);
dequeue(queue);
assert.isTrue(isEmpty(queue));
```

# --seed--

## --seed-contents--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  
}

function dequeue(queue) {
  
}

function front(queue) {
  
}

function size(queue) {
  
}

function isEmpty(queue) {
  
}
```

# --solutions--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  queue.collection.push(element);
}

function dequeue(queue) {
  return queue.collection.shift();
}

function front(queue) {
  return queue.collection[0];
}

function size(queue) {
  return queue.collection.length;
}

function isEmpty(queue) {
  return queue.collection.length === 0;
}
```
