---
id: 61fd6ab779390f49148773bb
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

Sotto il tuo elemento `h1`, crea un elemento `div`. Dagli un attributo `id` impostato su `years`. Vuoi che questo elemento in particolare sia nascosto ai lettori di schermo, quindi dagli l'attributo `aria-hidden` impostato su `true`.

# --hints--

Dovresti creare un nuovo elemento `div` dopo il tuo elemento `h1`.

```js
assert(document.querySelector('h1')?.nextElementSibling?.localName === 'div');
```

Il tuo nuovo elemento `div` dovrebbe avere un attributo `id` impostato su `years`.

```js
assert(document.querySelector('div')?.getAttribute('id') === 'years');
```

Il tuo nuovo elemento `div` dovrebbe avere l'attributo `aria-hidden` impostato su `true`.

```js
assert(document.querySelector('div')?.getAttribute('aria-hidden') === 'true');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
