---
id: 61fd5a93fd62bb35968adeab
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Prepara il tuo HTML con gli elementi `DOCTYPE`, `html` che indicano che questo documento è in inglese, `head` e `body`.

Dai al tuo elemento `head` gli appropriati elementi `meta` per `charset` e `viewport`, un elemento `title` con un titolo adeguato e un elemento `link` per il tuo foglio di stile.

# --hints--

Il tuo codice dovrebbe avere una dichiarazione `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Il tuo codice dovrebbe avere un elemento `html`.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

Il tuo elemento `html` dovrebbe avere un tag di apertura con un attributo `lang` di `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Il tuo elemento `html` dovrebbe avere un tag di chiusura.

```js
assert(code.match(/<\/html\s*>/));
```

Il tuo codice dovrebbe avere un elemento `head` all’interno dell’elemento `html`.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

Il tuo codice dovrebbe avere un elemento `body` all’interno dell’elemento `html`.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

Il tuo elemento `head` dovrebbe venire prima del tuo elemento `body`.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

Dovresti avere due elementi `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

Uno degli elementi `meta` dovrebbe avere un insieme `name` impostato su `viewport` e `content` impostato su `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

L’altro elemento `meta` dovrebbe avere l’attributo `charset` impostato su `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Il tuo codice dovrebbe avere un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Il tuo `title` dovrebbe contenere del testo.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

Il tuo codice dovrebbe avere un elemento `link`.

```js
assert(/<link/.test(code))
```

Il tuo elemento `link` dovrebbe essere all’interno del tuo elemento `head`.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i))
```

Il tuo elemento `link` dovrebbe avere un attributo `rel` con il valore `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

Il tuo elemento `link` dovrebbe avere un attributo `href` con il valore `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
