---
id: 61fd75ea7f663457612dba02
title: Passaggio 11
challengeType: 0
dashedName: step-11
---

# --description--

All’interno di ciascuno dei tuoi nuovi elementi `th`, annida un elemento `span` con l’attributo `class` impostato su `sr-only year`. Assegna loro il seguente testo (in ordine): `2019`, `2020` e `2021`.

Dai al tuo terzo elemento `th` l’attributo `class` impostato su `current`.

Lascia vuoto l’elemento `td`. Questo elemento serve solo a garantire che la tua tabella abbia un layout a quattro colonne e ad associare le intestazioni alle colonne corrette.

# --hints--

Ciascuno dei tuoi elementi `th` dovrebbe avere un elemento `span`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.length === 1);
  assert(th?.children?.[0]?.localName === 'span');
});
```

Ciascuno dei tuoi nuovi elementi `span` dovrebbe avere l’attributo `class` impostato su `sr-only year`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.[0]?.classList?.contains('sr-only'));
  assert(th?.children?.[0]?.classList?.contains('year'));
});
```

Il tuo primo elemento `span` dovrebbe contenere il testo `2019`.

```js
assert(document.querySelectorAll('th')?.[0]?.children?.[0]?.textContent === '2019');
```

Il tuo secondo elemento `span` dovrebbe contenere il testo `2020`.

```js
assert(document.querySelectorAll('th')?.[1]?.children?.[0]?.textContent === '2020');
```

Il tuo terzo elemento `span` dovrebbe contenere il testo `2021`.

```js
assert(document.querySelectorAll('th')?.[2]?.children?.[0]?.textContent === '2021');
```

Il tuo terzo elemento `th` dovrebbe avere l’attributo `class` impostato su `current`.

```js
assert(document.querySelector('table')?.querySelectorAll('th')?.[2]?.classList?.contains('current'));
```

Il tuo elemento `td` dovrebbe essere vuoto.

```js
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.textContent === '');
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.children?.length === 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
              <tr>
                <td></td>
                <th></th>
                <th></th>
                <th></th>
              </tr>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
