---
id: 662a6bc12cde72c32fb526f0
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

<dfn>インターフェース</dfn>はクラスの設計図のようなものです。インターフェースには、クラスが実装すべきメソッドやプロパティのセットが含まれています。

このプロジェクトは、`Equation`という名前の空のクラスを宣言することから始めてください。このクラスを使って、一般的な方程式の設計図となるインターフェースを定義します。

# --hints--

`Equation`という名前の新しいクラスを定義してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
