---
id: 662bde88dc84f1e249801b1a
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

`ABC` は抽象基底クラス（Abstract Base Classes）を表します。`ABC` クラスを使うと、通常のクラスを抽象クラスに変えることができます。抽象クラスは具体的なクラスの設計図として機能します。

`import` 文を修正して、`ABC` モジュールから `abc` クラスだけをインポートしてください。モジュール `x` から特定のオブジェクト `y` をインポートするには、インポート構文 `from y import x` を使います。

その後、`Equation` クラスを `ABC` を継承させて抽象クラスにしてください。

# --hints--

`ABC` を `abc` モジュールからインポートしてください。

```js
({ test: () => assert(runPython(`_Node(_code).has_import("from abc import ABC")`)) })
```

`Equation` クラスは `ABC` を継承する必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").inherits_from("ABC")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import abc


class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass
    
eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
