---
id: 662bdd364bf2cde1487922a9
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

他のプログラミング言語とは異なり、Pythonはコア言語でインターフェースを実装していませんが、Python標準ライブラリでは簡単にインターフェースを定義できます。

このプロジェクトでは、`abc` モジュールのユーティリティを使います。したがって、このモジュールをコードにインポートしてください。

# --hints--

`abc` モジュールをインポートしてください。

```js
({ test: () => assert(runPython(`_Node(_code).has_import("import abc")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
    

eq = Equation()
lin_eq = LinearEquation()
```
