---
id: 662bd456896f16d9bd03f1a6
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

`Equation` クラス内に、`solve` と `analyze` という名前の新しいインスタンスメソッドを2つ定義してください。

# --hints--

`solve` クラス内に `Equation` という名前のメソッドを定義してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("solve")`)) })
```

`solve` メソッドは `self` という1つのパラメータを受け取る必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_args("self")`)) })
```

`analyze` クラス内に `Equation` という名前のメソッドを定義してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("analyze")`)) })
```

`analyze` メソッドは `self` という1つのパラメータを受け取る必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Equation:
    pass
--fcc-editable-region--
```
