---
id: 662bd8260da84bdd5feae419
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

`LinearEquation` クラスには、`Equation` クラス内で定義されたすべてのメソッドを単に継承するのではなく、実装してほしいです。`Equation` クラスはインターフェースとして機能するべきです。

現在、`Equation` クラスは単に `LinearEquation` の親クラスになっています。次のステップで、それを正式なインターフェースに変える方法を学びます。

今は、`Equation` のインスタンスを作成して変数 `eq` に代入し、`LinearEquation` のインスタンスを作成して変数 `lin_eq` に代入してください。

# --hints--

変数 `eq` を宣言し、`Equation` のインスタンスを代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("eq = Equation()")`)) })
```

変数 `lin_eq` を宣言し、`LinearEquation` のインスタンスを代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation()")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
--fcc-editable-region--

--fcc-editable-region--
```
