---
id: 662bd552e1c1d2db1b88ba47
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

次に、`LinearEquation` という名前の別のクラスを定義し、`Equation` から継承させてください。このクラスは線形方程式を表すために使います。

# --hints--

`LinearEquation` という名前のクラスを定義してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_class("LinearEquation")`)) })
```

`LinearEquation` クラスは `Equation` クラスから継承する必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").inherits_from("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
--fcc-editable-region--

--fcc-editable-region--
```
