---
id: 662fa2e2cf27c09f21f4f5d0
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

Pythonでは、データ型は実行時（コードが実行されるとき）に認識されます。そのため、変数を宣言するときにデータ型を指定する必要はありません。それでも、変数が特定のデータ型を保持することを明確にするために、`variable: <data type> = value` または単に `variable: <data type>` を使って変数に注釈を付けることができます。Pythonインタープリターは変数に付けられた型注釈を強制しないことに注意してください。通常は外部ツールを使って型の検証を行います。

`Equation` クラスの中に、クラス属性 `degree` を定義してください。値は代入しないでください。代わりに、具体的なクラス内で整数を格納することを示すために、`int` の型注釈を使ってください。

後で、このクラス属性を使って、方程式オブジェクトをインスタンス化するときに渡される引数のバリデーションの一部として利用します。

# --hints--

`degree` クラス内に、`int` という名前のクラス属性を定義し、`Equation` で注釈を付ける必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_variable("degree").is_equivalent("degree: int")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    def __init__(self):
        pass
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass
--fcc-editable-region--
lin_eq = LinearEquation()
```
