---
id: 662fc3eba556a6bf800d48c1
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

`isinstance()` 組み込み関数は2つの引数を取り、最初の引数として渡されたオブジェクトが2番目の引数として渡されたクラスのインスタンスであるかどうかを示すブール値を返します。

```py
isinstance(7, int) # True
```

もう一つ確認したいのは、すべての引数が数値であることです。最初の `if` の後に、`for` を反復する `args` ループを作成し、現在の反復での引数が `int` または `float` のインスタンスでないかどうかをチェックしてください。`isinstance()` 関数を使い、2番目の引数として `int` と `float` を含むタプルを渡してください。

もし引数が数値でなければ、`TypeError` を発生させ、メッセージとして `"Coefficients must be of type 'int' or 'float'"` を使ってください。

# --hints--

`for` 文の後に `args` を反復する `if` ループを作成してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_iter().is_equivalent("args")`)) })
```

`if` ループ内で、現在の係数が `int` または `float` のいずれのインスタンスでもないかどうかをチェックする `for` 文を作成してください。

```js
({ test: () => assert(runPython(`
var = str(_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_vars())
cond1 = f'not isinstance({var}, (int, float))'
cond2 = f'not isinstance({var}, (float, int))'
if_stmt = _Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_conditions()[0]
if_stmt.is_equivalent(cond1) or if_stmt.is_equivalent(cond2)
`)) })
```

`TypeError` 文内で、提供された文字列を使って `if` を発生させてください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_bodies()[0].has_stmt("raise TypeError(\\"Coefficients must be of type 'int' or 'float'\\")")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
```
