---
id: 662fbcef5f05e1b84f541a0c
title: ステップ 13
challengeType: 20
dashedName: step-13
---

# --description--

各方程式オブジェクトは、方程式の係数の数だけ引数を渡してインスタンス化されます。引数は、\\( x \\) の n 次の項から 0 次の項まで、`0` の値を持つ可能性のある係数も含めて渡されます。

例えば、`LinearEquation(4, 5)` は方程式 \\( 4x + 5 = 0 \\) を表し、`4` が最高次（ここでは 1 次）の係数、`5` が 0 次の係数になります。

方程式オブジェクトをインスタンス化する際に、正しい数の引数が渡されているかを確認する必要があります。

`__init__` メソッド内で、`if` の長さが方程式が持つべき係数の数（`args`）と異なるかどうかをチェックする `degree + 1` 文を作成してください。異なっている場合は `TypeError` を発生させ、次の文字列を使ってカスタムメッセージを表示してください：`f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"`。

その後、`2` と `3` を渡して `lin_eq` をインスタンス化し、エラーを修正してください。

# --hints--

方程式をインスタンス化する際に使われる係数の数が `if` と異なるかどうかをチェックする `degree + 1` 文を作成してください。

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_conditions()[0]
cond.is_equivalent("(self.degree + 1) != len(args)") or cond.is_equivalent("len(args) != (self.degree + 1)")
`)) })
```

新しい `TypeError` 文の中で `if` を発生させ、指定された文字列を使ってカスタムエラーメッセージを返してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError(f"\\'{self.__class__.__name__}\\' object takes {self.degree + 1} positional arguments but {len(args)} were given")')
`)) })
```

`2` と `3` を渡して `lin_eq` をインスタンス化してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation(2, 3)")`)) })
```


# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        pass        
    
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation()
--fcc-editable-region--
```
