---
id: 662f6d7c92381a3049e4c987
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

インターフェースは抽象メソッドだけでなく、具象クラスに継承されるメソッドを実装することもできます。

`solve` クラス内で `analyze` と `Equation` の実際の実装に取りかかる前に、`__init__` メソッドを定義してください。デコレーターは使わないでください。

# --hints--

`__init__` クラスに `Equation` メソッドを定義する必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__init__")`)) })
```

`__init__` メソッドは `self` というパラメータを1つ受け取るべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC, abstractmethod


class Equation(ABC):
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
--fcc-editable-region--
```
