---
id: 6639fdcc701833a54c364211
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

係数を検証した後、それらをインスタンス属性にストアする必要があります。辞書内包表記を使って、キーが係数の次数で対応する値が係数となる辞書を作成し、`coefficients`という名前の属性に代入してください。

例えば、`LinearEquation`と`2`でインスタンス化された`4`オブジェクトは、`coefficients`属性として`{1: 2, 0: 4}`を持つべきです。なぜなら`2`が`x`の1次に対応し、`4`が`x`の0次に対応するからです。

新しい辞書のキーと値のペアは、`args`の順序に従って作成してください。

# --hints--

`coefficients`メソッド内で`__init__`という名前の属性を宣言してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_variable("self.coefficients")`)) })
```

係数をストアするために辞書内包表記を使うべきです。

```js
({ test: () => runPython(`
import ast
node = _Node(_code).find_class("Equation").find_function("__init__").find_variable("self.coefficients")
assert isinstance(node.tree.value, ast.DictComp)
`) })
```

`coefficients`属性は、次数-係数の形のキーと値のペアを含む辞書であるべきです。係数が`args`内に格納されている順序に従うことを忘れないでください。

```js
({ test: () => runPython(`
actual1 = list(LinearEquation(1, 6).coefficients.items())
expected1 = list({1: 1, 0: 6}.items())
actual2 = list(LinearEquation(-3.5, 0).coefficients.items())
expected2 = list({1: -3.5, 0: 0}.items())
assert actual1 == expected1
assert actual2 == expected2
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
--fcc-editable-region--
        
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass

    
lin_eq = LinearEquation(2, 3)
```
