---
id: 6639f947d3a1818c9322c64a
title: ステップ 16
challengeType: 20
dashedName: step-16
---

# --description--

係数の検証の最後のステップは、最高次数の係数がゼロでないことを確認することです。最高次数の係数は、オブジェクトをインスタンス化するときに最初の引数として渡されるべきであることを覚えておいてください。

そのための`if`文を追加し、次の文字列を使ってカスタムメッセージを提供する`ValueError`を発生させてください：`'Highest degree coefficient must be different from zero'`。

# --hints--

方程式をインスタンス化するときに渡された最初の係数がゼロかどうかをチェックする`if`文を作成してください。

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[2].find_conditions()[0]
cond.is_equivalent("args[0] == 0") or cond.is_equivalent("0 == args[0]") or cond.is_equivalent("not args[0]")
`)) })
```

新しい`ValueError`文の中で`if`を発生させ、提供された文字列を使ってカスタムエラーメッセージを返してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[2].find_bodies()[0].has_stmt("raise ValueError('Highest degree coefficient must be different from zero')")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
```
