---
id: 663b83a28943e6aa6275a514
title: ステップ 19
challengeType: 20
dashedName: step-19
---

# --description--

まだ`Equation`クラスの中で、作成する方程式オブジェクトに適切な文字列表現を与える`__str__`メソッドを定義してください。

今は、`__str__`メソッドの中で、変数`terms`を宣言し、空のリストを代入してください。この変数は方程式の各項（係数×\\( x^n \\)）をストアするために使います。

次に、変数`equation_string`を宣言し、`terms`リストの要素をスペースで結合した結果を代入してください。最後に`equation_string`を返してください。

# --hints--

`__str__`クラスの中に`Equation`メソッドを定義する必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__str__")`)) })
```

`__str__`メソッドは`self`というパラメータを受け取るべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_args("self")`)) })
```

`terms`メソッドの中で変数`__str__`を宣言し、空のリストを代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("terms = []")`)) })
```

`equation_string`メソッドの中で`terms`の要素をスペースで結合した結果を変数`__str__`に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("equation_string = ' '.join(terms)")`)) })
```

`equation_string`メソッドから`__str__`を返してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_return("equation_string")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    
--fcc-editable-region--    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
