---
id: 663c9f31306353460da54542
title: ステップ 23
challengeType: 20
dashedName: step-23
---

# --description--

ご覧の通り、出力に`+`記号が表示されていません。数値記号はデフォルトでは負の数の場合のみ表示されます。この動作を変更するには、f文字列の波括弧内で評価する式の後にコロンを書き、オプション`+`を指定してください。これにより、正の数と負の数の両方で記号を表示できます。

2つの条件節の文字列を修正し、`:+`の後の波括弧内に`coefficient`を追加してください。

# --hints--

`terms`文の中の`if`リストに追加する文字列を`f'{coefficient:+}'`に修正してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[1].find_bodies()[0].is_equivalent("terms.append(f'{coefficient:+}')")`)) })
```

`terms`節の中の`elif`リストに挿入する文字列を`f'{coefficient:+}x'`に修正してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[1].find_bodies()[1].is_equivalent("terms.append(f'{coefficient:+}x')")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient}')
            elif n == 1:
                terms.append(f'{coefficient}x')                
        equation_string = ' '.join(terms)        
        return equation_string
--fcc-editable-region--    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
