---
id: 664e4a590b52ba8d2adff19f
title: ステップ 30
challengeType: 20
dashedName: step-30
---

# --description--

二次方程式 \\( ax^2 + bx + c = 0 \\) の判別式は、通常ギリシャ文字の大文字デルタで表され、\\( Δ = b^2 - 4ac \\) と等しくなります。

`QuadraticEquation` クラス内に `__init__` メソッドを定義してください。親クラスの `super()` メソッドを呼び出すために `__init__` を使います。そして、方程式の判別式の値を格納する `delta` という新しい属性を定義してください。

# --hints--

`__init__` クラス内に `QuadraticEquation` メソッドを定義する必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").has_function("__init__")`)) })
```

`__init__` メソッドは `self` と `*args` の2つのパラメータを受け取るべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("__init__").has_args("self, *args")`)) })
```

`super().__init__(*args)` メソッド内で `__init__` を呼び出すべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("__init__").has_call("super().__init__(*args)")`)) })
```

`delta` メソッド内に `__init__` 属性を宣言し、方程式の判別式の値を代入するべきです。

```js
({ test: () => runPython(`
eq = QuadraticEquation(2, -3, -4)
assert eq.delta == 41
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}

class QuadraticEquation(Equation):
    degree = 2
--fcc-editable-region--
    
--fcc-editable-region--    
    def solve(self):
        pass
    
    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
print(lin_eq)

```
