---
id: 66759e32b88fb5459b1e0234
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

`__init_subclass__` メソッドは、そのメソッドを定義しているクラスがサブクラス化されるたびに呼び出され、子クラスをカスタマイズできるようにします。このメソッドは慣例で `cls`（「クラス」を表す）という名前のパラメータを受け取り、新しい子クラスを表します。

`__init_subclass__` クラスに `Equation` メソッドを定義し、`cls` パラメータを与えてください。

# --hints--

`__init_subclass__` クラスに `cls` パラメータを持つ `Equation` メソッドを定義してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").has_args("cls")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    
--fcc-editable-region--
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
