---
id: 6675aaf418b41157f6ccd692
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

`__init__` メソッドに戻る時が来ました。方程式のタイプによっては、インスタンス化の際に可変個の引数を渡す必要があります。

メソッドに2つ目のパラメータ `args` を追加し、`*` 演算子を使って可変個の引数を受け取れるようにしてください。

# --hints--

`__init__` メソッドは2つのパラメータ、`self` と `*args` を受け取る必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self, *args")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    degree = 1

    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
