---
id: 6899b593d232e07617f74055
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

在 `button` 元素內，創建一個包含 `😊` 的 `span` 元素。

給 `span` 元素一個值為 `img` 的 `role` 屬性和一個值為 `true` 的 `aria-hidden` 屬性。

然後，創建第二個 `span` 元素，其 `class` 為 `count`，包含文字 `0/10`。

# --hints--

在 `button` 元素內應該有兩個 `span` 元素。

```js
assert.lengthOf(document.querySelectorAll("button > span"), 2);
```

第一個 `span` 元素應該有 `😊` 的文字內容。

```js
assert.equal(document.querySelectorAll("button > span")[0]?.textContent.trim(), "😊");
```

第一個 `span` 元素應該有一個值為 `img` 的 `role` 屬性。

```js
assert.equal(document.querySelectorAll("button > span")[0]?.getAttribute("role"), "img");
```

第一個 `span` 元素應該有一個值為 `true` 的 `aria-hidden` 屬性。

```js
assert.equal(document.querySelectorAll("button > span")[0]?.getAttribute("aria-hidden"), "true");
```

第二個 `span` 元素應該有一個值為 `count` 的 `class` 屬性。

```js
assert.equal(document.querySelectorAll("button > span")[1]?.className, "count");
```

第二個 `span` 元素應該有 `0/10` 的文字內容。

```js
assert.equal(document.querySelectorAll("button > span")[1]?.textContent.trim(), "0/10");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Emoji Reactor</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <main>
    <h1 class="title">How are you feeling today?</h1>
    <p class="description">
      Click on the buttons below to rate your emotions.
    </p>
    <div class="btn-container">
      <button id="happy-btn" class="emoji-btn" aria-label="Happy face emoji">
--fcc-editable-region--
        
--fcc-editable-region--
      </button>
    </div>
  </main>
</body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --light-grey: #efefef;
  --white: #ffffff;
  --very-dark-blue: #0a0a23;
  --light-purple: #a78bfa;
  --very-light-purple: #c4b5fd;
  --purple: #8b5cf6;
}

body {
  background-color: var(--very-dark-blue);
  color: var(--light-grey);
  font-family: sans-serif;
}

main {
  text-align: center;
  padding: 10px;
}

.btn-container,
button {
  display: flex;
  justify-content: space-evenly;
  align-items: center;
}

.btn-container {
  flex-direction: column;
}

.emoji-btn {
  width: 70%;
  cursor: pointer;
  color: var(--white);
  background-color: var(--light-purple);
  background-image: linear-gradient(
    to bottom,
    var(--very-light-purple),
    var(--light-purple)
  );
  border: 3px solid var(--purple);
  border-radius: 8px;
  padding: 10px;
  font-size: 1.5rem;
  margin: 10px 0;
  transition: background-color 0.3s ease, transform 0.2s ease;
}

@media (min-width: 768px) {
  .emoji-btn {
    width: 30%;
  }
}

.emoji-btn:hover {
  background-color: var(--purple);
  background-image: none;
}

.title {
  margin-top: 15px;
  font-size: 2rem;
}

.description {
  font-size: 1.4rem;
  margin: 20px 0;
}
```
