---
id: 5f331e55dfab7a896e53c3a1
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

The `title` is one of several elements that provide extra information not visible on the web page, but it is useful for search engines or how the page gets displayed.

在`head`元素中，嵌套一个 `meta` 元素，其属性名为 `charset`，值为`utf-8`，以告诉浏览器如何对页面进行编码。

# --hints--

你应该有一个 `meta` 标签。

```js
assert.match(code, /<meta.*?\/?>/is);
```

`meta` 元素是一个空元素，它不应该有一个结束标签 `</meta>`。

```js
assert.notMatch(code, /<\/meta>/i);
```

Your `meta` tag should have a `charset` attribute.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Your `charset` attribute should have a value of `utf-8`.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
