---
id: 5f3326b143638ee1a09ff1e3
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

To prepare to create some actual content, add a `body` element below the `head` element.

# --hints--

你应该有一个 `<body>` 开始标签。

```js
assert(code.match(/<body>/i));
```

你应该有一个 `</body>` 结束标签。

```js
assert(code.match(/<\/body>/i));
```

你不应该改变你的 `head` 元素。 确认你没有删除结束标签。

```js
assert(code.match(/<head>/i));
assert(code.match(/<\/head>/i));
```

你的 `body` 元素应该在你的 `head` 元素之后。

```js
assert(code.match(/<\/head>[.\n\s]*<body>/im));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```

