---
id: 587d8259367417b2b2512c85
title: Implementar o Algoritmo de Ordenação por Seleção
challengeType: 26
dashedName: implement-the-selection-sort-algorithm
---

# --description--

Neste laboratório você vai implementar o algoritmo de ordenação por seleção. A ordenação por seleção funciona selecionando o menor valor em uma lista e trocando-o com o primeiro valor da lista. Depois, começa na segunda posição, seleciona o menor valor na lista restante e troca-o com o segundo elemento. Continua iterando pela lista e trocando elementos até chegar ao final da lista. Agora a lista está ordenada. A ordenação por seleção tem complexidade de tempo quadrática em todos os casos.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter uma `selectionSort` função.
1. A `selectionSort` função deve receber um array e retornar um array com os mesmos elementos, mas ordenados do menor para o maior.

# --hints--

`selectionSort` deve ser uma função.

```js
assert.isFunction(selectionSort);
```

`selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` deve retornar um array que não foi alterado exceto pela ordem.

```js
assert.sameMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`selectionSort` deve retornar um array ordenado (do menor para o maior elemento).

```js
assert.sameOrderedMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
)
```

`selectionSort` não deve usar o método `.sort()` integrado.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
