---
id: 6443c98ac169e501eba36d04
title: 23단계
challengeType: 0
dashedName: step-23
---

# --description--

`timeAgo` 함수에서는 현재 시간과 주제에서 마지막 활동이 있었던 시간 사이의 차이를 계산해야 합니다. 이렇게 하면 주제에 활동이 있은 후 얼마나 시간이 지났는지 표시할 수 있습니다.

다음 요구사항을 충족하는 `timeAgo` 함수를 완성하세요:

- 경과한 분 수가 `60` 미만이면 문자열 `"xm ago"`을 반환하세요. `x`는 분을 나타냅니다.
- 경과한 시간 수가 `24` 미만이면 문자열 `"xh ago"`을 반환하세요. `x`은 시간을 나타냅니다.
- 그렇지 않으면 문자열 `"xd ago"`을 반환하세요. `x`는 일을 나타냅니다.

시간 차이를 계산하는 데 도움이 되는 몇 가지 식은 다음과 같습니다:

- `minutes = Math.floor((currentTime - lastPost) / 60000);`
- `hours = Math.floor((currentTime - lastPost) / 3600000);`
- `days = Math.floor((currentTime - lastPost) / 86400000);`

# --hints--

경과한 분 수가 `"30m ago"`분일 때 문자열 `30`을 반환해야 합니다.

```js
const currentTime = new Date();
const lastPostTime = new Date(currentTime.getTime() - 30 * 60 * 1000); 
const result = timeAgo(lastPostTime.toISOString());

assert.strictEqual(result, '30m ago');
```

경과한 시간 수가 `"1h ago"`시간일 때 문자열 `1`을 반환해야 합니다.

```js
const currentTime = new Date();
const lastPostTime = new Date(currentTime.getTime() - 60 * 60 * 1000);
const result = timeAgo(lastPostTime.toISOString());

assert.strictEqual(result, '1h ago');
```

경과한 일 수가 `"1d ago"`일일 때 문자열 `1`을 반환해야 합니다.

```js
const currentTime = new Date();
const lastPostTime = new Date(currentTime.getTime() - 24 * 60 * 60 * 1000);
const result = timeAgo(lastPostTime.toISOString());

assert.strictEqual(result, '1d ago');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>fCC Forum Leaderboard</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <nav>
        <img
          class="fcc-logo"
          src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
          alt="freeCodeCamp logo"
        />
      </nav>
      <h1 class="title">Latest Topics</h1>
    </header>
    <main>
      <div class="table-wrapper">
        <table>
          <thead>
            <tr>
              <th id="topics">Topics</th>
              <th id="avatars">Avatars</th>
              <th id="replies">Replies</th>
              <th id="views">Views</th>
              <th id="activity">Activity</th>
            </tr>
          </thead>
          <tbody id="posts-container"></tbody>
        </table>
      </div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

.fcc-logo {
  width: 210px;
  display: block;
  margin: auto;
}

.title {
  margin: 25px 0;
  text-align: center;
  color: var(--light-grey);
}

.table-wrapper {
  padding: 0 25px;
  overflow-x: auto;
}

table {
  width: 100%;
  color: var(--dark-grey);
  margin: auto;
  table-layout: fixed;
  border-collapse: collapse;
  overflow-x: scroll;
}

#topics {
  text-align: start;
  width: 60%;
}

th {
  border-bottom: 2px solid var(--dark-grey);
  padding-bottom: 10px;
  font-size: 1.3rem;
}

td:not(:first-child) {
  text-align: center;
}

td {
  border-bottom: 1px solid var(--dark-grey);
  padding: 20px 0;
}

.post-title {
  font-size: 1.2rem;
  color: var(--medium-grey);
  text-decoration: none;
}

.category {
  padding: 3px;
  color: var(--black);
  text-decoration: none;
  display: block;
  width: fit-content;
  margin: 10px 0 10px;
}

.career {
  background-color: var(--salmon-color);
}

.feedback,
.html-css {
  background-color: var(--light-blue);
}

.support {
  background-color: var(--light-orange);
}

.general {
  background-color: var(--light-green);
}

.javascript {
  background-color: var(--golden-yellow);
}

.backend {
  background-color: var(--gold);
}

.python {
  background-color: var(--green);
}

.motivation {
  background-color: var(--peach);
}

.avatar-container {
  display: flex;
  justify-content: center;
  gap: 10px;
  flex-wrap: wrap;
}

.avatar-container img {
  width: 30px;
  height: 30px;
}

@media (max-width: 750px) {
  .table-wrapper {
    padding: 0 15px;
  }

  table {
    width: 700px;
  }

  th {
    font-size: 1.2rem;
  }

  .post-title {
    font-size: 1.1rem;
  }
}
```

```js
const forumLatest = "https://cdn.freecodecamp.org/curriculum/forum-latest/latest.json";
const forumTopicUrl = "https://forum.freecodecamp.org/t/";
const forumCategoryUrl = "https://forum.freecodecamp.org/c/";
const avatarUrl = "https://sea1.discourse-cdn.com/freecodecamp";

const postsContainer = document.getElementById("posts-container");

--fcc-editable-region--
const timeAgo = (time) => {
  const currentTime = new Date();
  const lastPost = new Date(time);
  
  
};
--fcc-editable-region--

const fetchData = async () => {
  try {
    const res = await fetch(forumLatest);
    const data = await res.json();
    showLatestPosts(data);
  } catch (err) {
    console.log(err);
  }
};

fetchData();

const showLatestPosts = (data) => {
  const { topic_list, users } = data;
  const { topics } = topic_list;

  postsContainer.innerHTML = topics.map((item) => {
    const {
      id,
      title,
      views,
      posts_count,
      slug,
      posters,
      category_id,
      bumped_at,
    } = item;

    return `
    <tr>
      <td>
        <p class="post-title">${title}</p>
      </td>
      <td></td>
      <td>${posts_count - 1}</td>
      <td>${views}</td>
      <td></td>
    </tr>`;
  }).join("");
};


```
