---
id: 68e97fe79367ad7b5dd6c9cd
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

На цьому кроці додайте елемент `h1` з текстом `XYZ Bookstore`.

# --hints--

У вас має бути елемент `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Текст вашого елемента `h1` має бути `XYZ Bookstore`. Перевірте орфографію та правильність використання великих літер.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'XYZ Bookstore');
```

Ваш елемент `h1` має бути всередині елемента `body`.

```js
assert.match(code, /<body>[\s\S]*<h1>[\s\S]*<\/h1>[\s\S]*<\/body>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
