---
id: 68ebdbacdd3fa474132cc975
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Під елементом `h1` додайте елемент `p` з таким текстом: `Browse our collection of amazing books!`.

# --hints--

У вас має бути елемент `p`.

```js
assert.exists(document.querySelector('p'));
```

Текст вашого елемента `p` має бути: `Browse our collection of amazing books!`. 

```js
assert.equal(document.querySelector('p')?.innerText.trim(), 'Browse our collection of amazing books!');
```

Ваш елемент `p` має бути розташований під вашим елементом `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
