---
id: 68eca6364c5616d393389a30
title: Крок 18
challengeType: 0
dashedName: step-18
---

# --description--

Пам’ятайте, HTML-елемент виглядає так:

```html
<element attribute="value">
    inner text
</element>
```

Під елементом з класом `card-container` додайте новий елемент `p` з таким текстом:

```md
Review your selections and continue to checkout.
```

Під елементом `p` створіть елемент `div` з атрибутом `class`, встановленим у `btn-container`. Цей контейнер згрупує ваші елементи кнопок навігації.

# --hints--

Новий доданий елемент `p` має бути розташований під елементом з класом `card-container`.

```js
assert.exists(document.querySelector('.card-container + p'));
```

Текст вашого нового елемента `p` має бути `Review your selections and continue to checkout.`

```js
const cardContainer = document.querySelector('.card-container');
assert.equal(cardContainer?.nextElementSibling?.innerText.trim(), 'Review your selections and continue to checkout.');
```

У вас має бути елемент `div` з атрибутом `class`, значення якого `btn-container`.

```js
assert.exists(document.querySelector('div.btn-container'));
```

Ваш елемент `div` з атрибутом `class`, значення якого `btn-container`, має йти після новододаного елемента `p`.

```js
assert.exists(document.querySelector('p + .btn-container'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
