---
id: 68ec99e478211dc578699944
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Під елементом `h2` у першому елементі з класом `card` додайте елемент `p` з таким текстом:

```md
This is an epic story of Sally and her dog Rex as they navigate through other worlds.
```

# --hints--

У вас має бути вкладений елемент `p` всередині елемента з класом `card`.

```js
assert.exists(document.querySelector('.card p'));
```

Текст вашого елемента `p` має бути `This is an epic story of Sally and her dog Rex as they navigate through other worlds.`

```js
assert.equal(document.querySelector('.card p')?.innerText.trim(), "This is an epic story of Sally and her dog Rex as they navigate through other worlds.");
```

Ваш елемент `p` має бути розташований нижче за елемент `h2`.

```js
assert.exists(document.querySelector('.card h2 + p'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
