---
id: 68ec6d9a315221aa31e54816
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Атрибут `class` використовується для ідентифікації одного або кількох елементів для стилізації. На відміну від атрибута `id`, імена класів не повинні бути унікальними: кілька елементів можуть мати однаковий клас.

Ось приклад:

```html
<p class="example">example paragraph</p>
```

Додайте атрибут `class` до вашого елемента `div` і встановіть його значення на `card-container`.

# --hints--

Ваш елемент `div` має мати атрибут `class`.

```js
assert.isTrue(document.querySelector('div')?.hasAttribute('class'));
```

Ваш елемент `div` має мати атрибут `class` зі значенням `card-container`.

```js
assert.equal(document.querySelector('div')?.className, 'card-container');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  <div>
--fcc-editable-region--
    
  </div>
</body>

</html>
```
