---
id: 68eca2a795b333ca5fee30a8
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Тепер створіть другу картку книги. Додайте ще один елемент `div` з атрибутом `class`, встановленим у `card`. Зверніть увагу, як можна повторно використовувати однакову назву класу для кількох елементів, щоб застосувати послідовне стилювання.

# --hints--

У вас має бути два елементи `div` з атрибутом `class`, значення якого дорівнює `card`.

```js
assert.lengthOf(document.querySelectorAll('.card'), 2);
```

Другий елемент із класом `card` має бути всередині елемента з класом `card-container`.

```js
assert.lengthOf(document.querySelectorAll('.card-container .card'), 2);
```

Два елементи з класом `card` мають бути братськими.

```js
assert.exists(document.querySelector('.card + .card'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
  --fcc-editable-region--
    
  --fcc-editable-region--
  </div>
</body>

</html>
```
