---
id: 68ec98b38d83a3c28dd30efe
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Всередині першого елемента з класом `card` додайте елемент `h2` з текстом `Sally's SciFi Adventure`.

# --hints--

У вас має бути вкладений елемент `h2` всередині елемента з класом `card`.

```js
assert.exists(document.querySelector('.card h2'));
```

Текст вашого елемента `h2` має бути `Sally's SciFi Adventure`. 

```js
assert.equal(document.querySelector('.card h2')?.innerText.trim(), "Sally's SciFi Adventure");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
