---
id: 68ec9332a9b5b2b32487bd00
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Атрибут `id` додає унікальний ідентифікатор до HTML-елемента. Кожен `id` має бути унікальним на сторінці і використовуватися лише один раз.

Значення `id` не можуть містити пробілів і повинні містити лише літери, цифри, підкреслення та дефіси.

Ось приклад:

```html
<p id="para">example paragraph</p>
```

Додайте атрибут `id` до вашого елемента з класом `card` і встановіть його значення на `sally-adventure-book`.

# --hints--

Ваш елемент з класом `card` має мати атрибут `id`.

```js
assert.isTrue(document.querySelector('.card')?.hasAttribute('id'));
```

Ваш елемент з класом `card` має мати `id` зі значенням `sally-adventure-book`.

```js
assert.equal(document.querySelector('.card')?.id, 'sally-adventure-book');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
  --fcc-editable-region--
    <div class="card">
  --fcc-editable-region--
      
    </div>
  </div>
</body>

</html>
```
