---
id: 68ec6cd5b7e8f5a8f7319e32
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Елемент `div` використовується як контейнер для групування інших HTML-елементів. Ви переважно будете використовувати елемент `div`, коли хочете згрупувати HTML-елементи, які матимуть спільну множину CSS-стилів.

Під елементом `p` додайте елемент `div`. Цей `div` буде контейнером для ваших карток книг.

# --hints--

У вас має бути елемент `div`.

```js
assert.exists(document.querySelector('div'));
```

Ваш елемент `div` має бути розташований нижче за елемент `p`.

```js
assert.exists(document.querySelector('p + div'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
