---
id: 68eca11e99e3c5c894ca9d69
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Елемент `button` використовується для створення активних кнопок на вебсторінці. Кнопки — це інтерактивні елементи, на які користувачі можуть натиснути, щоб виконати дії.

Додайте елемент `button` всередину елемента, який має `class` зі значенням `card`, надайте кнопці атрибут `class` зі значенням `btn` і текст `Buy Now`.

# --hints--

У вас має бути вкладений елемент `button` всередині елемента з класом `card`.

```js
assert.exists(document.querySelector('.card button'));
```

Ваш елемент `button` повинен мати атрибут `class`.

```js
assert.isTrue(document.querySelector('.card button')?.hasAttribute('class'));
```

Ваш елемент `button` повинен мати `class` зі значенням `btn`.

```js
assert.equal(document.querySelector('.card button')?.className, 'btn');
```

Текст вашого елемента `button` має бути `Buy Now`.

```js
assert.equal(document.querySelector('.card button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
