---
id: 68eca5c0065c82d256d29ca3
title: Крок 17
challengeType: 0
dashedName: step-17
---

# --description--

Всередині другої картки додайте елемент `button` з атрибутом `class`, встановленим у `btn`, та текстом `Buy Now`.

Обидва елементи `button` тепер мають однаковий `class`, що означає, що їх можна стилізувати послідовно разом.

# --hints--

У другому елементі, який має клас `button`, має бути вкладений елемент `card`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('button'));
```

Елемент `button` у другій картці повинен мати атрибут `class`.

```js
const cards = document.querySelectorAll('.card');
assert.isTrue(cards[1]?.querySelector('button')?.hasAttribute('class'));
```

Елемент `button` у другій картці повинен мати `class` зі значенням `btn`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.className, 'btn');
```

Текст елемента `button` у другій картці має бути `Buy Now`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
