---
id: 69134b6cb43aa5f254950a10
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті ви створите сторінку книжкового магазину, створюючи картки книг, які відображають інформацію про різні книги. Ви попрактикуєтеся в організації вмісту за допомогою `div` елементів, класів і ідентифікаторів.

Почніть свою сторінку книжкового магазину зі створення шаблонного коду HTML.

Додайте оголошення `<!DOCTYPE html>` та елементи `html` і `head`.

Додайте атрибут `lang` до елемента `html` і встановіть його в `"en"`.

# --hints--

Ваш код має містити оголошення `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Ваш код має містити як відкриваючий, так і закриваючий тег `<html>`.

```js
assert.match(code,/<html\b[^>]*>[\s\S]*<\/html>/i);
```

Ваш елемент `html` має мати атрибут `lang`, встановлений у `"en"`.

```js
assert.match(code, /<html[^>]*\blang\s*=\s*["']en["'][^>]*>/i);
```

Ваш код має містити як відкриваючий, так і закриваючий тег `<head>`.

```js
assert.match(code,/<head>[\s\S]*<\/head>/i);
```

Ваш елемент `head` має бути вкладений у елемент `html`.

```js
assert.match(code, /<\/html>\s*$/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
