---
id: 69134d14150d51f8066a3246
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Тепер покращте структуру свого HTML-документа, щоб переконатися, що ваша сторінка правильно кодується.

Всередині елемента `head` додайте елемент `<meta charset="UTF-8">`.

Нарешті, додайте елемент `body` нижче розділу `head`. Саме тут розміщуватиметься весь видимий вміст сторінки.

# --hints--

Ваш код має містити елемент `<meta charset="UTF-8">`.

```js
assert.exists(document.querySelector('head > meta[charset="utf-8"]'));
```

Ваш код має містити відкриваючий і закриваючий теги `<body>`.

```js
assert.match(code,/<body>[\s\S]*<\/body>/i);
```

Ваш елемент `body` має бути розташований нижче елемента `head`.

```js
assert.match(code, /<\/head>[\s\S]*<body>[\s\S]*<\/body>[\s\S]*<\/html>/i);
```

# --seed--
## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
  <title>XYZ Bookstore Page</title>
</head>

--fcc-editable-region--
</html>
```
