---
id: 691f85085787a8c4ca589fbb
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Додайте елемент `title` всередину елемента `head`.

Встановіть заголовок сторінки `XYZ Bookstore Page`.

# --hints--

Ваш код має містити як відкриваючий, так і закриваючий теги `<title>`.

```js
assert.match(code,/<title>[\s\S]*<\/title>/i);
```

Ваш заголовок має бути `XYZ Bookstore Page`.

```js
assert.equal(document.querySelector('title')?.innerText.trim(),'XYZ Bookstore Page');
```

Ваш елемент `title` має бути вкладений у елемент `head`.

```js
assert.match(code, /<head>[\s\S]*<title>[\s\S]*<\/title>[\s\S]*<\/head>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
--fcc-editable-region-- 
</head>

</html>
```
