---
id: 66a207974c806a19d6607073
title: Створіть застосунок для інвентаризації книг
challengeType: 25
dashedName: build-a-book-inventory-app
demoType: onClick
---

# --description--


Виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Ви повинні мати елемент `h1` з текстом `Book Inventory`.
1. Ви повинні мати елемент `table` зі стовпцями під назвою `Title`, `Author`, `Category`, `Status` та `Rate`.
1. Таблиця повинна мати принаймні чотири рядки: в першому мають бути заголовки стовпців, а в інших — інформація про книги.
1. Кожен рядок з тіла таблиці повинен мати клас `read`, `to-read` або `in-progress`.
1. Елементи `td` стовпця `Status` мають містити елемент `span` з `class` зі значенням `status`, що оточує текст `Read`, `To Read` або `In Progress` залежно від класу рядка.
1. Елементи `td` стовпця `Rate` мають містити елемент `span` з `class` зі значенням `rate`, що охоплює три порожні елементи `span`.
1. Елементи `.rate`, які розміщені в рядках `.read`, повинні мати додатковий клас зі значенням `one`, `two` або `three`, що залежить від особистої оцінки. Це значення має бути після `rate`.
1. Створіть три селектори атрибутів, щоб націлитись на елементи з класом `read`, `to-read` та `in-progress`, і встановіть їхню властивість `background-image` на `linear-gradient` за власним вибором.
1. Встановіть значення властивості `display` всіх елементів `span` на `inline-block`.
1. Використайте селектор атрибутів, щоб націлитись на елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `to-read`, і встановіть їхні властивості `border` та `background-image`.
1. Використайте селектор атрибутів, щоб націлитись на елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `read`, і встановіть їхні властивості `border` та `background-image`.
1. Використайте селектор атрибутів, щоб націлитись на елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `in-progress`, і встановіть їхні властивості `border` та `background-image`.
1. Використайте селектор атрибутів, щоб націлитись на елементи `span` з класом `status` і елементи `span`, значення класу яких починається з `rate`, та встановіть їхні властивості `height`, `width`, `padding`.
1. Використайте селектор атрибутів, щоб націлитись на елементи `span`, які є прямими дочірніми для елементів `span`, значення `class` яких починається з `rate`, та встановіть їхні властивості `border`, `border-radius`, `margin`, `height`, `width`, `background-color`.
1. Використайте селектор атрибутів, щоб націлитись на першого нащадка елементів `span`, значення `class` яких містить слово `one`, і встановіть його властивість `background-image` на `linear-gradient`.
1. Використайте селектор атрибутів, щоб націлитись на перших двох нащадків елементів `span`, значення `class` яких містить слово `two`, і встановіть їхню властивість `background-image` на `linear-gradient`.
1. Використайте селектор атрибутів, щоб націлитись на трьох елементів `span`, які є нащадками елементів `span`, значення `class` яких містить слово `three`, і встановіть їхню властивість `background-image` на `linear-gradient`.

# --hints--

Ви повинні мати елемент `h1` з текстом `Book Inventory`.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Book Inventory');
```

Ви повинні мати тільки один елемент `h1`.

```js
assert.equal(document.querySelectorAll('h1')?.length, 1);
```

You should have a `table` element.

```js
assert.equal(document.querySelectorAll('table')?.length, 1);
```

Елемент `table` повинен мати п’ять стовпців.

```js
assert.equal(document.querySelectorAll('th')?.length, 5);
```

Перший стовпець повинен мати заголовок з текстом `Title`.

```js
assert.equal(document.querySelectorAll('th')[0]?.innerText, 'Title');
```

Другий стовпець повинен мати заголовок з текстом `Author`.

```js
assert.equal(document.querySelectorAll('th')[1]?.innerText, 'Author');
```

Третій стовпець повинен мати заголовок з текстом `Category`.

```js
assert.equal(document.querySelectorAll('th')[2]?.innerText, 'Category');
```

Четвертий стовпець повинен мати заголовок з текстом `Status`.

```js
assert.equal(document.querySelectorAll('th')[3]?.innerText, 'Status');
```

П’ятий стовпець повинен мати заголовок з текстом `Rate`.

```js
assert.equal(document.querySelectorAll('th')[4]?.innerText, 'Rate');
```

Таблиця повинна містити принаймні чотири рядки.

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
```

Кожен рядок таблиці (окрім заголовка) повинен мати клас `read`, `to-read` або `in-progress`.

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
for (let i = 1; i < rows.length; i++) {
  const classList = [...rows[i].classList];
  const currentProgress = classList[0]; 
  assert.oneOf(currentProgress, ['read', 'to-read', "in-progress"]); 
}
```

Елементи `td` стовпця `Status` мають містити елемент `span`.

```js
const statusColumnData = document.querySelectorAll('td:nth-child(4)');
assert.isAbove(statusColumnData.length, 0);
for (let e of statusColumnData) {
  assert.equal(e?.children[0]?.tagName, 'SPAN');
}
```

Елементи `span` стовпця `Status` повинні мати клас `status`.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
assert.isAbove(statusSpans.length, 0);
for (let e of statusSpans) {
    assert.isTrue(e.classList.contains('status'));
}
```

Кожен елемент `.status` повинен містити текст `Read`, `To Read` або `In Progress` залежно від класу його рядка.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
const rows = Array.from(document.querySelectorAll('tr')).slice(1);
assert.isAbove(statusSpans.length, 0);
for (let i = 0; i < rows.length; i++) {
  switch (statusSpans[i]?.innerText) {
        case 'Read':
            assert.isTrue(rows[i].classList.contains('read'));
            break;
        case 'To Read':
            assert.isTrue(rows[i].classList.contains('to-read'));
            break;
        case 'In Progress':
            assert.isTrue(rows[i].classList.contains('in-progress'));
            break;
        default:
            assert.fail(); 
    }
}
```

Елементи `td` стовпця `Rate` мають містити елемент `span`.

```js
const rateColumnData = document.querySelectorAll('tr td:last-child');
assert.isAbove(rateColumnData.length, 0);
for (let e of rateColumnData) {
    assert.equal(e.children[0]?.tagName, 'SPAN')
}
```

Першим класом кожного елемента `span` — який є прямим дочірнім для елемента `td` стовпця `Rate` — має бути `rate`.

```js
const rateSpans = document.querySelectorAll('tr td:last-child > span:first-child');
assert.isAbove(rateSpans.length, 0);
for (let e of rateSpans) {
    assert.equal(e.classList[0], 'rate');
}
```

Кожен елемент `.rate`повинен містити три порожні елементи `span`.

```js
const rateSpans = document.getElementsByClassName('rate');
assert.isAbove(rateSpans.length, 0);
for (let e of rateSpans) {
    assert.equal(e.children.length, 3);
    for (let child of e.children) {
        assert.equal(child.tagName, 'SPAN');
        assert.equal(child.innerText.length, 0);
    }
}
```

Елементи `.rate`, які розміщені в рядках `.read`, після класу `rate` повинні мати ще один клас зі значенням `one`, `two` або `three`, що залежить від особистої оцінки.

```js
const readBooksRates = document.querySelectorAll('.read .rate');
assert.isAbove(readBooksRates.length, 0);
for (let e of readBooksRates) {
    assert.oneOf(e.classList[1], ['one', 'two', 'three']);
}
```

Ви повинні мати селектор атрибутів, щоб націлитись на рядки з класом `read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"]'));
```

Використайте селектор атрибутів, щоб націлитись на рядки з класом `read`, і встановіть їхню властивість `background-image` на лінійний градієнт за власним вибором.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="read"]')?.backgroundImage, 'linear-gradient(');
```

Ви повинні мати селектор атрибутів, щоб націлитись на рядки з класом `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]'));
```

Використайте селектор атрибутів, щоб націлитись на рядки з класом `to-read`, і встановіть їхню властивість `background-image` на лінійний градієнт за власним вибором.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]')?.backgroundImage, 'linear-gradient(');
```

Ви повинні мати селектор атрибутів, щоб націлитись на рядки з класом `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]'));
```

Використайте селектор атрибутів, щоб націлитись на рядки з класом `in-progress`, і встановіть їхню властивість `background-image` на лінійний градієнт за власним вибором.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]')?.backgroundImage, 'linear-gradient(');
```

Встановіть значення властивості `display` всіх елементів `span` на `inline-block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('span')?.getPropVal('display'), 'inline-block');
```

Ви повинні мати селектор атрибутів, щоб націлитись на елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]'));
```

Використайте селектор атрибутів, щоб націлитись на елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `to-read`, і встановіть їхню властивість `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.border);
```

Використайте селектор атрибутів, щоб націлитись на елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `to-read`, і встановіть їхню властивість `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.backgroundImage);
```

Ви повинні мати селектор атрибутів, щоб націлитись на елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]'));
```

Використайте селектор атрибутів, щоб націлитись на елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `read`, і встановіть їхню властивість `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.border);
```

Використайте селектор атрибутів, щоб націлитись на елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `read`, і встановіть їхню властивість `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.backgroundImage);
```

Ви повинні мати селектор атрибутів, щоб націлитись на елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]'));
```

Використайте селектор атрибутів, щоб націлитись на елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `in-progress`, і встановіть їхню властивість `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.border);
```

Використайте селектор атрибутів, щоб націлитись на елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `in-progress`, і встановіть їхню властивість `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.backgroundImage);
```

Ви повинні мати селектор атрибутів, щоб націлитись на елементи `span` з класом `status` і елементи `span`, значення класу яких починається з `rate`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.exists(selector1 || selector2);
```

Використайте селектор атрибутів, щоб націлитись на елементи `span` з класом `status` і елементи `span`, значення класу яких починається з `rate`, та встановіть їхню властивість `height`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.height || selector2?.height);
```

Використайте селектор атрибутів, щоб націлитись на елементи `span` з класом `status` і елементи `span`, значення класу яких починається з `rate`, та встановіть їхню властивість `width`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.width || selector2?.width);
```

Використайте селектор атрибутів, щоб націлитись на елементи `span` з класом `status` і елементи `span`, значення класу яких починається з `rate`, та встановіть їхню властивість `padding`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.padding || selector2?.padding);
```

Ви повинні мати селектор атрибутів, щоб націлитись на елементи `span`, які є дочірніми для елементів `span`, значення `class` яких починається з `rate`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span'));
```

Використайте селектор атрибутів, щоб націлитись на елементи `span`, які є прямими дочірніми для елементів `span`, значення `class` яких починається з `rate`, та встановіть їхню властивість `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border'));
```

Використайте селектор атрибутів, щоб націлитись на елементи `span`, які є прямими дочірніми для елементів `span`, значення `class` яких починається з `rate`, та встановіть їхню властивість `border-radius`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border-radius'));
```

Використайте селектор атрибутів, щоб націлитись на елементи `span`, які є прямими дочірніми для елементів `span`, значення `class` яких починається з `rate`, та встановіть їхню властивість `margin`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('margin'));
```

Використайте селектор атрибутів, щоб націлитись на елементи `span`, які є прямими дочірніми для елементів `span`, значення `class` яких починається з `rate`, та встановіть їхню властивість `height`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('height'));
```

Використайте селектор атрибутів, щоб націлитись на елементи `span`, які є прямими дочірніми для елементів `span`, значення `class` яких починається з `rate`, та встановіть їхню властивість `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('width'));
```

Використайте селектор атрибутів, щоб націлитись на елементи `span`, які є прямими дочірніми для елементів `span`, значення `class` яких починається з `rate`, та встановіть їхню властивість `background-color`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('background-color'));
```

Ви повинні мати селектор атрибутів, щоб націлитись на першого нащадка елементів `span`, значення `class` яких містить слово `one`.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Використайте селектор атрибутів, щоб націлитись на першого нащадка елементів `span`, значення `class` яких містить слово `one`, і встановіть його властивість `background-image` на `linear-gradient`.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image').includes('linear-gradient('));
```

Ви повинні мати селектор атрибутів, щоб націлитись на перших двох нащадків елементів `span`, значення `class` яких містить слово `two`.

```js

const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Використайте селектор атрибутів, щоб націлитись на перших двох нащадків елементів `span`, значення `class` яких містить слово `two`, і встановіть їхню властивість `background-image` на `linear-gradient`.

```js
const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]

assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.backgroundImage.includes('linear-gradient(')) 
```

Ви повинні мати селектор атрибутів, щоб націлитись на елементи `span`, які є нащадками елементів `span`, значення `class` яких містить слово `three`.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.exists(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Використайте селектор атрибутів, щоб націлитись на елементи `span`, які є нащадками елементів `span`, значення `class` яких містить слово `three`, і встановіть їхню властивість `background-image` на `linear-gradient`.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.include(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image'), 'linear-gradient(');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main>
    <h1>Book Inventory</h1>
    <div class="table-container">
      <table>
        <thead>
          <tr>
            <th>Title</th>
            <th>Author</th>
            <th>Category</th>
            <th>Status</th>
            <th>Rate</th>
          </tr>
        </thead>
        <tbody>
          <tr class="read">
            <td>The Three Musketeers</td>
            <td>Alexandre Dumas</td>
            <td>Historical Novel</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate three">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="to-read">
            <td>The Plague</td>
            <td>Albert Camus</td>
            <td>Philosophical Novel</td>
            <td><span class="status">To Read</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>The Metamorphosis</td>
            <td>Franz Kafka</td>
            <td>Novella</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate one">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Dead Souls</td>
            <td>Nikolai Gogol</td>
            <td>Picaresque</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="in-progress">
            <td>Lord of the Flies</td>
            <td>William Golding</td>
            <td>Allegorical Novel</td>
            <td><span class="status">In Progress</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Do Androids Dream of Electric Sheep?</td>
            <td>Philip K. Dick</td>
            <td>Science Fiction</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
        </tbody>
      </table>
    </div>
  </main>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
    font-family: Arial, sans-serif;
}

.table-container {
  overflow-x: auto;
}

table {
    border-collapse: collapse;
    width: 100vw;
}

thead {
    background-image: linear-gradient(hsl(197, 92%, 77%, 0.60), hsl(197, 92%, 50%, 0.60));
}

th {
    padding: 0.4em;
}

td {
    text-align: center;
    padding: 0.3em;
}

tr[class="read"] {
    background-image: linear-gradient(hsl(120, 100%, 85%, 0.6), hsl(120, 100%, 65%, 0.6));
}

tr[class="to-read"] {
    background-image: linear-gradient(hsl(120, 20%, 95%, 0.6), hsl(120, 20%, 75%, 0.6));
}

tr[class="in-progress"] {
    background-image: linear-gradient(hsl(40, 100%, 85%, 0.60), hsl(40, 100%, 65%, 0.6));
}

span {
    display: inline-block;
}

span[class="status"] {
    border-radius: 20%/60%;
}

span[class="status"],
span[class^="rate"] {
    height: 1.85em;
    width: 7em;
    padding: 0.3em;
}

tr[class="to-read"] span[class="status"] {
    border: 0.1em solid hsl(5, 100%, 45%);
    background-image: linear-gradient(hsl(5, 100%, 75%), hsl(5, 100%, 50%));
}

tr[class="read"] span[class="status"] {
    border: 0.1em solid hsl(120, 75%, 45%);
    background-image: linear-gradient(hsl(120, 75%, 75%), hsl(120, 75%, 50%));
}

tr[class="in-progress"] span[class="status"] {
    border: 0.1em solid hsl(40, 90%, 40%);
    background-image: linear-gradient(hsl(40, 90%, 75%), hsl(40, 90%, 50%));
}

span[class^="rate"] > span {
    border: 0.1em solid hsl(0, 0%, 50%);
    border-radius: 50%;
    background-color: hsl(0, 15%, 95%);
    height: 100%;
    width: 20%;
    margin: 0.1em;
}

span[class~="three"] span {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="two"] :nth-child(1),
span[class~="two"] :nth-child(2) {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="one"] :first-child {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}
```
