---
id: 6564aee9c077774ea49c3faf
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

El algoritmo de ordenamiento por mezcla realiza tres acciones:

- Dividir una secuencia desordenada de elementos en subpartes
- Ordenar los elementos en las subpartes
- Unir las subpartes ordenadas

Lo anterior sucede recursivamente hasta que las subpartes se combinan en una secuencia completamente ordenada. Comencemos dividiendo la secuencia.

Primero, reemplaza la palabra clave `pass` con una variable `middle_point`. Utiliza el operador de división entera (`//`) para dividir la longitud de la lista `array` entre 2 y asigna el resultado a tu nueva variable `middle_point`. Recuerda sangrar tu código.

# --hints--

No debes tener `pass` en el cuerpo de la función.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).find_function("merge_sort").has_pass()`))
})
```

Debes declarar una variable `middle_point` y asignarle `len(array) // 2`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_stmt("middle_point = len(array) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
