---
id: 656639771fed09b5c0e8fe71
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Ahora que has dividido la lista de `array` en dos listas separadas, continuarás dividiendo cada lista hasta que cada elemento esté solo en su propia lista. Una lista con un solo número siempre está ordenada.

Para hacer eso, llama recursivamente a `merge_sort` dentro de tu función y pasa `left_part` como argumento a la llamada.

# --hints--

Debes llamar a `merge_sort` con el argumento `left_part` al final del cuerpo de tu función.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("right_part = array[middle_point:]", "merge_sort(left_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
