---
id: 656680b0fc79f2c38a34d90e
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Ahora es el momento de ordenar y fusionar las listas (`left_part` y `right_part`) en la lista original.

Puedes hacer esto comparando elementos en ambas listas y fusionando el elemento más pequeño en la lista principal. Vas a hacer esta comparación para todos los índices en `left_part` y `right_part`.

Crea tres variables: `left_array_index`, `right_array_index` y `sorted_index` y establece sus valores en `0`. Estas variables te ayudarán a realizar un seguimiento de cada índice durante el proceso de ordenamiento.

# --hints--

Debes tener una variable llamada `left_array_index` dentro de tu función `merge_sort`

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_array_index")`)) 

})
```

El valor de `left_array_index` debe ser `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("left_array_index").is_equivalent("left_array_index = 0")`))  

})
```

Debes tener una variable llamada `right_array_index` dentro de tu función `merge_sort`

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_array_index")`))
})
```

El valor de `right_array_index` debe ser `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_array_index").is_equivalent("right_array_index = 0")`))
})

```

Debes tener una variable llamada `sorted_index` dentro de tu función `merge_sort`

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("sorted_index")`))
})
```

El valor de `sorted_index` debe ser `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("sorted_index").is_equivalent("sorted_index = 0")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)
--fcc-editable-region--

--fcc-editable-region--
```
