---
id: 656709e50ed928da35260276
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

En un paso anterior, asignaste el elemento en el `left_part` a la lista `array` porque era más pequeño. Sin embargo, esto no siempre será el caso. En algunos casos de comparación, el elemento de la derecha podría ser más pequeño.

Crea una cláusula `else` para ejecutar cuando el elemento en `left_part` no sea menor que el elemento en `right_part`.

Dentro del bloque `else`, asigna `right_part[right_array_index]` a `array[sorted_index]`.

# --hints--

Deberías tener `array[sorted_index] = right_part[right_array_index]` en el cuerpo de tu cláusula `else`.

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_bodies()[1].find_body()[0].is_equivalent("array[sorted_index] = right_part[right_array_index]")
`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
--fcc-editable-region--
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
--fcc-editable-region--
```
