---
id: 656707e11e671ed54c96f7ec
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

Cuando la `if` condición se evalúa como `True`, significa que el elemento en la lista `left_part` es menor que el elemento al que se está comparando en la lista `right_part`.

En ese caso, puedes colocar el elemento encontrado en `left_array_index` dentro de `left_part` en el arreglo ordenado.

Dentro del bloque `if`, elimina `pass` y asigna `left_part[left_array_index]` a `array[sorted_index]`.

# --hints--

Debes sustituir `pass` con `array[sorted_index] = left_part[left_array_index]` en el cuerpo de la instrucción `if`.

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_body().is_equivalent("array[sorted_index] = left_part[left_array_index]")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            pass

--fcc-editable-region--
```
