---
id: 65670d1ef177e7e2b76d9528
title: Paso 14
challengeType: 20
dashedName: step-14
---

# --description--

El `if` y `else` que creaste en los pasos anteriores asignarán elementos al array ordenado.

Cada elemento asignado al array ordenado ocupa un índice en la lista. Por lo tanto, debes moverte al siguiente índice en el array ordenado después de cada asignación.

Debajo del bloque `if`/`else`, pero aún dentro del bucle `while`, aumenta `sorted_index` por `1`.

# --hints--

Debes incrementar el valor actual de `sorted_index` por `1` después de tu cláusula `else`.

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_bodies()[0].find_body()[1]
    cond.is_equivalent("sorted_index += 1") or cond.is_equivalent("sorted_index = sorted_index + 1") or cond.is_equivalent("sorted_index = 1 + sorted_index")
    `))
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
--fcc-editable-region--
```
