---
id: 6569b743630ee592a65a7e2f
title: Paso 20
challengeType: 20
dashedName: step-20
---

# --description--

Antes de probar la función `merge_sort()`, necesitas crear un caso base que detenga la ejecución de la función cuando la longitud de `array` sea menor o igual a `1`.

Este caso base detendrá la llamada de recursión. Sin él, la operación de ordenación por mezcla continuaría ejecutándose incluso cuando la lista ya haya sido ordenada o no tenga elementos.

Justo después de la declaración de la función, crea una declaración `if` con esta condición: `len(array) <= 1`. Dentro del cuerpo de la declaración `if`, agrega una instrucción `return` para detener la ejecución de la función.

# --hints--

Debes crear una declaración `if` para verificar si `len(array) <= 1`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_function("merge_sort").find_ifs()[0].find_conditions()[0].is_equivalent("len(array) <= 1")
    `))
})
```

Debes tener una declaración `return` en el cuerpo de la declaración `if`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_ifs()[0].find_body().has_stmt("return")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):

--fcc-editable-region--

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1

    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1
```
