---
id: 6569b4e0bd29d17d4b53b16c
title: Paso 19
challengeType: 20
dashedName: step-19
---

# --description--

Ahora, vas a replicar la misma lógica del bucle `while` para `right_part`.

Crea un bucle `while` que se ejecute cuando `right_array_index` sea menor que `len(right_part)`. Dentro del bucle `while`, asigna `right_part[right_array_index]` a `array[sorted_index]`. Al final, incrementa `right_array_index` y `sorted_index` en `1`.

# --hints--

Tu bucle `while` debe verificar si `right_array_index` es menor que `len(right_part)`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_conditions()[0].is_equivalent("right_array_index < len(right_part)")`)) 
})
```

Debes asignar `right_part[right_array_index]` a `array[sorted_index]` en el cuerpo del bucle.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_body().has_stmt("array[sorted_index] = right_part[right_array_index]")`))
})
```

Debes incrementar `right_array_index` en `1` dentro del bucle `while`. Ten en cuenta que se debe realizar el incremento después de asignar el elemento de la parte derecha al arreglo ordenado.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()

    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = right_array_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = 1 + right_array_index")
    `))
}) 
```

Debes incrementar el valor actual de `sorted_index` en `1` dentro del bucle `while`. Ten en cuenta que el incremento debe hacerse después de asignar el elemento de la parte derecha al arreglo ordenado.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()

    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = 1 + sorted_index")
    `))
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1

--fcc-editable-region--

--fcc-editable-region--
```
