---
id: 6577254891048c8f2c19e961
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Al final de tu funci贸n, llama a la funci贸n `merge_sort()` nuevamente. Esta vez, pasa `right_part` como el argumento en la llamada de la funci贸n.

# --hints--

Debes llamar a `merge_sort()` con el argumento `right_part` al final de tu funci贸n.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("merge_sort(left_part)", "merge_sort(right_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)

--fcc-editable-region--
```
