---
id: 657f3a661730891aa64f3568
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Usa la sintaxis slice para extraer la mitad derecha de `array` y asignarla a una variable llamada `right_part`.

# --hints--

Debe tener una variable llamada `right_part`

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_part")`))
})
```

Debe usar la sintaxis slice para extraer la mitad derecha del array y asignarla a la variable `right_part`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_part").is_equivalent("right_part = array[middle_point:]")`)) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
--fcc-editable-region--
```
