---
id: 587d78b0367417b2b2512b05
title: 기술 문서 페이지 만들기
challengeType: 14
saveSubmissionToDB: true
forumTopicId: 301146
dashedName: build-a-technical-documentation-page
---

# --description--

**목표:** <a href="https://technical-documentation-page.freecodecamp.rocks" target="_blank" rel="noopener noreferrer nofollow">https://technical-documentation-page.freecodecamp.rocks</a>와 기능적으로 유사한 앱을 만드세요. **이 데모 프로젝트를 복사하지 마세요**.

**사용자 스토리:**

1. `main` 요소와 해당하는 `id="main-doc"`가 보이며, 이 요소에는 페이지의 주요 내용(기술 문서)이 포함되어 있습니다.
1. `#main-doc` 요소 안에는 `section` 요소가 여러 개 있으며, 각각 `main-section` 클래스를 가지고 있습니다. 최소 다섯 개가 있어야 합니다.
1. 각 `.main-section` 안의 첫 번째 요소는 `header` 요소여야 하며, 이 요소에는 해당 섹션의 주제를 설명하는 텍스트가 포함되어야 합니다.
1. `main-section` 클래스를 가진 각 `section` 요소에는 `id`가 있어야 하며, 이 `id`는 그 안에 포함된 각 `header` 텍스트와 일치해야 합니다. 공백은 밑줄(_)로 교체되어야 합니다(예: "JavaScript and Java"라는 헤더가 있는 섹션은 해당하는 `id="JavaScript_and_Java"`를 가져야 합니다).
1. `.main-section` 요소에는 총 최소 10개의 `p` 요소가 포함되어야 합니다(각각이 아니라 전체 합계).
1. `.main-section` 요소에는 총 최소 5개의 `code` 요소가 포함되어야 합니다(각각이 아니라 전체 합계).
1. `.main-section` 요소에는 총 최소 5개의 `li` 항목이 포함되어야 합니다(각각이 아니라 전체 합계).
1. `nav` 요소와 해당하는 `id="navbar"`가 보여야 합니다.
1. 내비게이션 바 요소에는 기술 문서 주제를 설명하는 텍스트가 포함된 `header` 요소가 하나 있어야 합니다.
1. 또한 내비게이션 바에는 `nav-link` 클래스를 가진 링크(`a`) 요소가 있어야 하며, `main-section` 클래스를 가진 각 요소마다 하나씩 있어야 합니다.
1. `#navbar` 안의 `header` 요소는 내비게이션 바의 링크(`a`) 요소들보다 앞에 있어야 합니다.
1. `nav-link` 클래스를 가진 각 요소는 각 `section` 안의 `header` 텍스트와 일치하는 텍스트를 포함해야 합니다(예: "Hello world" 섹션/헤더가 있다면 내비게이션 바에도 "Hello world" 텍스트가 포함된 요소가 있어야 합니다).
1. 내비게이션 바 요소를 클릭하면 페이지가 `#main-doc` 요소의 해당 섹션으로 이동해야 합니다(예: "Hello world" 텍스트가 있는 `.nav-link` 요소를 클릭하면, 해당 id를 가진 `section` 요소로 이동하고 그에 맞는 헤더가 포함되어야 합니다).
1. 일반 크기 기기(노트북, 데스크톱)에서는 `id="navbar"` 요소가 화면 왼쪽에 표시되어 항상 사용자에게 보여야 합니다.
1. 기술 문서에는 최소 하나 이상의 미디어 쿼리를 사용해야 합니다.

사용자 스토리를 모두 충족하고 아래 모든 테스트를 통과하면 이 프로젝트를 완료할 수 있습니다. 자신만의 스타일로 만드세요. 즐거운 코딩 되세요!

**참고:** HTML에 `<link rel="stylesheet" href="styles.css">`를 추가하여 스타일시트를 연결하고 CSS를 적용하세요.

# --hints--

`main` 요소가 `id`가 `main-doc`인 것을 포함해야 합니다.

```js
const el = document.getElementById('main-doc');
assert.isNotNull(el);
```

`main-section` 클래스를 가진 최소 다섯 개의 `section` 요소가 있어야 합니다.

```js
const els = document.querySelectorAll('#main-doc section');
assert.isAtLeast(els.length, 5);
```

모든 `.main-section` 요소는 `section` 요소여야 합니다.

```js
const els = document.querySelectorAll('.main-section');
els.forEach(el => {
  if (el.tagName !== 'SECTION') {
    assert.fail();
  }
});
assert.isAbove(els.length, 0);
```

`#main-doc`의 자손 요소인 최소 다섯 개의 `.main-section` 요소가 있어야 합니다.

```js
const els = document.querySelectorAll('#main-doc .main-section');
assert.isAtLeast(els.length, 5);
```

각 `.main-section`의 첫 번째 자식은 `header` 요소여야 합니다.

```js
const els = document.querySelectorAll('.main-section');
els.forEach(el => {
  if (el.firstElementChild?.tagName !== 'HEADER') assert.fail();
});
assert.isNotEmpty(els);
```

어떤 `header` 요소도 비어 있으면 안 됩니다.

```js
const els = document.querySelectorAll('header');
els.forEach(el => {
  if (el.innerText?.length <= 0) assert.fail();
});
assert.isNotEmpty(els);
```

모든 `.main-section` 요소는 `id`를 가져야 합니다.

```js
const els = document.querySelectorAll('.main-section');
els.forEach(el => {
  if (!el.id || el.id === '') assert.fail();
});
assert.isNotEmpty(els);
```

각 `.main-section`는 첫 번째 자식 텍스트와 일치하는 `id`를 가져야 하며, 자식 텍스트의 공백은 밑줄(_)로 교체되어야 합니다(`_`).

```js
const els = document.querySelectorAll('.main-section');
els.forEach(el => {
  const text = el.firstElementChild?.innerText?.replaceAll(' ', '_');
  if (el.id?.toUpperCase() !== text?.toUpperCase()) assert.fail();
});
assert.isNotEmpty(els);
```

`.main-section` 요소 안에 총 최소 10개의 `p` 요소가 있어야 합니다.

```js
const els = document.querySelectorAll('.main-section p');
assert.isAtLeast(els.length, 10);
```

`.main-section` 요소의 자손인 최소 다섯 개의 `code` 요소가 있어야 합니다.

```js
const els = document.querySelectorAll('.main-section code');
assert.isAtLeast(els.length, 5);
```

`.main-section` 요소의 자손인 최소 다섯 개의 `li` 요소가 있어야 합니다.

```js
const els = document.querySelectorAll('.main-section li');
assert.isAtLeast(els.length, 5);
```

`nav` 요소가 `id`가 `navbar`인 것을 포함해야 합니다.

```js
const el = document.getElementById('navbar');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'NAV');
```

`#navbar` 안에는 정확히 하나의 `header` 요소가 있어야 합니다.

```js
const els = document.querySelectorAll('#navbar header');
assert.lengthOf(els, 1);
```

`nav-link` 클래스를 가진 최소 하나의 `a` 요소가 있어야 합니다.

```js
const els = document.querySelectorAll('a.nav-link');
assert.isAtLeast(els.length, 1);
```

모든 `.nav-link` 요소는 앵커(`a`) 요소여야 합니다.

```js
const els = document.querySelectorAll('.nav-link');
els.forEach(el => {
  if (el.tagName !== 'A') assert.fail();
});
assert.isNotEmpty(els);
```

모든 `.nav-link` 요소는 `#navbar` 안에 있어야 합니다.

```js
const els1 = document.querySelectorAll('.nav-link');
const els2 = document.querySelectorAll('#navbar .nav-link');
assert.isNotEmpty(els2);
assert.strictEqual(els1.length, els2.length);
```

`.nav-link` 요소와 `.main-section` 요소의 개수가 같아야 합니다.

```js
const els1 = document.querySelectorAll('.main-section');
const els2 = document.querySelectorAll('.nav-link');
assert.isNotEmpty(els1);
assert.isNotEmpty(els2);
assert.strictEqual(els1.length, els2.length);
```

`#navbar` 안의 `header` 요소는 `#navbar` 안의 링크(`a`) 요소들보다 앞에 있어야 합니다.

```js
const navLinks = document.querySelectorAll('#navbar a.nav-link');
const header = document.querySelector('#navbar header');
navLinks.forEach(navLink => {
  if (
    header.compareDocumentPosition(navLink) & Node.DOCUMENT_POSITION_PRECEDING
  )
    assert.fail();
});
assert.isNotNull(header);
```

각 `.nav-link`는 관련 `section`의 `header` 텍스트와 일치하는 텍스트를 가져야 합니다(예: "Hello world" 섹션/헤더가 있다면 `#navbar`는 "Hello world" 텍스트를 가진 `.nav-link`를 가져야 합니다).

```js
const headerText = Array.from(document.querySelectorAll('.main-section')).map(
  el => el.firstElementChild?.innerText?.trim().toUpperCase()
);
const linkText = Array.from(document.querySelectorAll('.nav-link')).map(el =>
  el.innerText?.trim().toUpperCase()
);
const remainder = headerText.filter(str => linkText.indexOf(str) === -1);

assert.isNotEmpty(headerText);
assert.isNotEmpty(linkText);
assert.isEmpty(remainder);
```

각 `.nav-link`는 해당 `.main-section`에 연결하는 `href` 속성을 가져야 합니다(예: "Hello world" 텍스트가 있는 `.nav-link` 요소를 클릭하면, 해당 id를 가진 `section` 요소로 이동합니다).

```js
const hrefValues = Array.from(document.querySelectorAll('.nav-link')).map(el =>
  el.getAttribute('href')
);
const mainSectionIDs = Array.from(
  document.querySelectorAll('.main-section')
).map(el => el.id);
const missingHrefValues = mainSectionIDs.filter(
  str => hrefValues.indexOf('#' + str) === -1
);
assert.isNotEmpty(hrefValues);
assert.isNotEmpty(mainSectionIDs);
assert.isEmpty(missingHrefValues, 0);
```

`#navbar`는 항상 창의 왼쪽 가장자리에 있어야 합니다.

```js
const el = document.getElementById('navbar');
const left1 = el?.offsetLeft;
assert.isNotNull(el);
assert.isAtLeast(left1, -15);
assert.isAtMost(left1, 15);
```

기술 문서 프로젝트에는 최소 하나 이상의 미디어 쿼리를 사용해야 합니다.

```js
const htmlSourceAttr = Array.from(document.querySelectorAll('source')).map(el =>
  el.getAttribute('media')
);
const cssCheck = new __helpers.CSSHelp(document).getCSSRules('media');
assert.isTrue(cssCheck.length > 0 || htmlSourceAttr.length > 0);
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
    <title>Technical Documentation Page</title>
  </head>
  <body>
    <nav id="navbar">
      <header><br />Algebraic Concepts</header>
      <hr />
      <a href="#introduction" class="nav-link">Introduction</a><br />
      <hr />
      <a href="#definitions" class="nav-link">Definitions</a><br />
      <hr />
      <a href="#examples" class="nav-link">Examples</a><br />
      <hr />
      <a href="#solving_equations" class="nav-link">Solving Equations</a><br />
      <hr />
      <a href="#solving_equations_ii" class="nav-link">Solving Equations II</a
      ><br />
      <hr />
      <a href="#solving_equations_iii" class="nav-link">Solving Equations III</a
      ><br />
      <hr />
      <a href="#system_of_equations" class="nav-link">System of Equations</a
      ><br />
      <hr />
      <a href="#try_it_yourself!" class="nav-link">Try it Yourself!</a><br />
      <hr />
      <a href="#more_information" class="nav-link">More Information</a><br />
    </nav>
    <main id="main-doc">
      <section class="main-section" id="introduction">
        <header>Introduction</header>
        <p>
          Welcome to a basic introduction of algebra. In this tutorial, we will
          review some of the more common algebraic concepts.
        </p>
      </section>
      <section class="main-section" id="definitions">
        <header>Definitions</header>
        <p>
          To start with, let's define some of the more common terms used in
          algebra:
        </p>
        <ul>
          <li>
            <b>Variable:</b> A variable is an unknown value, usually represented
            by a letter.
          </li>
          <li>
            <b>Expression:</b> Essentially a mathematical object. For the
            purpose of this tutorial, an expression is one part of an equation.
          </li>
          <li>
            <b>Equation:</b> An equation is a mathematical argument in which two
            expressions result in the same value.
          </li>
        </ul>
      </section>
      <section class="main-section" id="examples">
        <header>Examples</header>
        <p>
          Sometimes it is easier to understand the definitions when you have a
          physical example to look at. Here is an example of the above terms.<br /><br />
          <code>x + 5 = 12 </code><br /><br />
          In this above example, we have:
        </p>
        <ul>
          <li><b>Variable:</b> The variable in the example is "x".</li>
          <li>
            <b>Expression:</b> There are two expressions in this example. They
            are "x+5" and "12".
          </li>
          <li>
            <b>Equation:</b> The entire example, "x+5=12", is an equation.
          </li>
        </ul>
      </section>
      <section class="main-section" id="solving_equations">
        <header>Solving Equations</header>
        <p>
          The primary use for algebra is to determine an unknown value, the
          "variable", with the information provided. Continuing to use our
          example from above, we can find the value of the variable "x".<br /><br />
          <code>x + 5 = 12 </code><br /><br />
          In an equation, both sides result in the same value. So you can
          manipulate the two expressions however you need, as long as you
          perform the same operation (or change) to each side. You do this
          because the goal when solving an equation is to
          <b
            >get the variable into its own expression, or by itself on one side
            of the = sign.</b
          ><br />For this example, we want to remove the "+5" so the "x" is
          alone. To do this, we can <em>subtract 5</em>, because subtraction is
          the opposite operation to addition. But remember, we have to perform
          the same operation to both sides of the equation. Now our equation
          looks like this.<br /><br />
          <code>x + 5 - 5 = 12 - 5</code><br /><br />
          The equation looks like a mess right now, because we haven't completed
          the operations. We can <b>simplify</b> this equation to make it easier
          to read by performing the operations "5-5" and "12-5". The result
          is:<br /><br />
          <code>x = 7</code><br /><br />
          We now have our solution to this equation!
        </p>
      </section>
      <section class="main-section" id="solving_equations_ii">
        <header>Solving Equations II</header>
        <p>
          Let us look at a slightly more challenging equation.<br /><br />
          <code>3x + 4 = 13</code><br /><br />
          Again we can start with subtraction. In this case, we want to subtract
          4 from each side of the equation. We will also go ahead and simplify
          with each step. So now we have:<br /><br />
          <code>3x = 9</code><br /><br />
          "3x" translates to "3*x", where the "*" symbol indicates
          multiplication. We use the "*" to avoid confusion, as the "x" is now a
          variable instead of a multiplication symbol. The opposite operation
          for multiplication is division, so we need to
          <b>divide each expression by 3</b>.<br /><br />
          <code>x = 3</code><br /><br />
          And now we have our solution!
        </p>
      </section>
      <section class="main-section" id="solving_equations_iii">
        <header>Solving Equations III</header>
        <p>
          Now we are getting in to more complex operations. Here is another
          equation for us to look at:<br /><br />
          <code>x^2 - 8 = 8</code><br /><br />
          Our very first step will be to <b>add</b> 8 to each side. This is
          different from our previous examples, where we had to subtract. But
          remember, our goal is to get the variable alone by performing opposite
          operations.<br /><br />
          <code>x^2 = 16</code><br /><br />
          But what does the "^2" mean? The "^" symbol is used to denote
          exponents in situations where superscript is not available. When
          superscript <b>is</b> available, you would see it as x<sup>2</sup>.
          For the sake of this project, however, we will use the "^" symbol.<br />
          An exponent tells you how many times the base (in our case, "x") is
          multiplied by itself. So, "x^2" would be the same as "x*x". Now the
          opposite function of multiplication is division, but we would have to
          <b>divide both sides by "x"</b>. We do not want to do this, as that
          would put an "x" on the other side of the equation. So instead, we
          need to use the root operation! For an exponent of "2", we call this
          the "square root" and denote it with "√". Our equation is now:
          <br /><br />
          <code>x = √9</code><br /><br />
          Performing a root operation by hand can be a tedious process, so we
          recommend using a calculator when necessary. However, we are lucky in
          that "9" is a
          <a href="https://en.wikipedia.org/wiki/Square_number"
            >perfect square</a
          >, so we do not need to calculate anything. Instead, we find our
          answer to be:<br /><br />
          <code>x = 3</code>
        </p>
      </section>
      <section class="main-section" id="system_of_equations">
        <header>System of Equations</header>
        <p>
          As you explore your algebra studies further, you may start to run
          across equations with more than one variable. The first such equations
          will likely look like:<br /><br />
          <code>y = 3x</code><br /><br />
          An equation like this does <b>not have one single solution</b>.
          Rather, there are a series of values for which the equation is true.
          For example, if "x=3" and "y=9", the equation is true. These equations
          are usually used to plot a graph. <br />
          Getting more complicated, though, you may be given a <b>pair</b> of
          equations. This is called a "system of equations", and CAN be solved.
          Let's look at how we do this! Consider the following system of
          equations:<br /><br />
          <code>y = 3x | y - 6 = x</code>
          A system of equations IS solvable, but it is a multi-step process. To
          get started, we need to choose a variable we are solving for. Let's
          solve for "x" first. From the second equation, we know that "x" equals
          "y - 6", but we cannot simplify that further because we do not have a
          value for "y". Except, thanks to the system of equations, we DO have a
          value for "y". We know that "y" equals "3x". So, looking at our second
          equation, we can replace "y" with "3x" because they have the same
          value. We then get:<br /><br />
          <code>3x - 6 = x</code><br /><br />
          Now we can solve for "x"! We start by adding 6 to each side.<br /><br />
          <code>3x = x + 6</code><br /><br />
          We still need to get "x" by itself, so we subtract "x" from both sides
          and get:<br /><br />
          <code>2x = 6</code><br /><br />
          If this confuses you, remember that "3x" is the same as "x+x+x".
          Subtract an "x" from that and you get "x+x", or "2x". Now we divide
          both sides by 2 and have our value for x!<br /><br />
          <code>x = 3</code><br /><br />
          However, our work is not done yet. We still need to find the value for
          "y". Let's go back to our first equation:<br /><br />
          <code>y = 3x</code><br /><br />
          We have a value for "x" now, so let's see what happens if we put that
          value in.<br /><br />
          <code>y = 3*3</code><br /><br />
          We perform the multiplication and discover that "y=9"! Our solution to
          this system of equations then is:<br /><br />
          <code>x = 3 and y = 9</code><br /><br />
        </p>
      </section>
      <section class="main-section" id="try_it_yourself!">
        <header>Try it Yourself!</header>
        <p>Coming Soon!</p>
        <p>Keep an eye out for new additions!</p>
      </section>
      <section class="main-section" id="more_information">
        <header>More Information</header>
        <p>Check out the following links for more information!</p>
        <ul>
          <li>
            <a href="https://www.wolframalpha.com/examples/mathematics/algebra/"
              >Wolfram Alpha</a
            >
            is a great source for multiple mathematic fields.
          </li>
          <li>
            <a href="https://en.wikipedia.org/wiki/Algebra"
              >Wikipedia's Algebra page</a
            >
            for more general information.
          </li>
        </ul>
      </section>
    </main>
  </body>
  <footer>
    <a href="../">Return to Project List</a> |
    <a href="https://www.nhcarrigan.com">Return to HomePage</a>
  </footer>
</html>
```

```css
* {
  background-color: #3a3240;
}
a {
  color: #92869c;
}
a:hover {
  background-color: #92869c;
  color: #3a3240;
}
#navbar {
  border-style: solid;
  border-width: 5px;
  border-color: #92869c;
  height: 100%;
  top: -5px;
  left: -5px;
  padding: 5px;
  text-align: center;
  color: #92869c;
}
@media (min-width: 480px) {
  #navbar {
    position: fixed;
  }
}
main {
  margin-left: 220px;
  color: #92869c;
}
header {
  font-size: 20pt;
}
code {
  background-color: #92869c;
  border-style: dashed;
  border-width: 2px;
  border-color: #92869c;
  padding: 5px;
  color: black;
}
footer {
  text-align: center;
}
```
