---
id: afcc8d540bea9ea2669306b6
title: 建置一個字串重複函式（程式）
challengeType: 26
dashedName: build-a-string-repeating-function
---

# --description--

在本實驗中，你將創建一個函式（程式），用來重複給定的字串指定的次數。為了本實驗的目的，請*不要*使用內建的 `.repeat()` 方法。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該創建一個名為 `repeatStringNumTimes` 的函式（程式），該函式接受兩個參數：一個字串和一個數字。
2. 該函式（程式）應傳回指定次數重複的字串。
3. 如果數字小於或相等於零，該函式應該傳回空字串。

# --hints--

你應該創建一個名為 `repeatStringNumTimes` 的函式（程式）。

```js
assert.isFunction(repeatStringNumTimes);
```

`repeatStringNumTimes` 應該接受兩個參數。

```js
assert.lengthOf(repeatStringNumTimes, 2);
```

函式（程式）`repeatStringNumTimes` 應該永遠傳回字串。

```js
assert.isString(repeatStringNumTimes('hello', 3));
```

`repeatStringNumTimes("*", 3)` 應該傳回字串 `***`。

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

`repeatStringNumTimes("abc", 3)` 應該傳回字串 `abcabcabc`。

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

`repeatStringNumTimes("abc", 4)` 應該傳回字串 `abcabcabcabc`。

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

`repeatStringNumTimes("abc", 1)` 應該傳回字串 `abc`。

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

`repeatStringNumTimes("*", 8)` 應該傳回字串 `********`。

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

`repeatStringNumTimes("abc", -2)` 應該傳回一個空字串（`""`）。

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

`repeatStringNumTimes("abc", 0)` 應該傳回 `""`。

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

內建的 `repeat()` 方法不應該被使用。

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
