---
id: 699d1f8e51fcdc8230f779fb
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utajenga mfumo wa orodha ya maktaba ya urithi. Utachambua data ghafi ya vitabu, kutafuta na kuunganisha maingizo, kuthibitisha ubora wa data, na kusafirisha matokeo kwa miundo tofauti.

Kadi kila moja ya orodha huhifadhiwa kama mfuatano wa herufi mmoja katika muundo `"Title | Author | Year | Location"`, na sehemu zimegawanywa na alama ya bomba (`|`).

Utaitumia safu ya data ya mfuatano huu wa herufi katika warsha nzima kujenga na kujaribu vitendakazi vya orodha yako. Tengeneza safu ya data iitwayo `rawCatalogCards` na ongeza mfuatano wa herufi mbili zifuatazo ndani yake:

```js
"From a Buick 8 | King, Stephen | 2002 | Shelf K7",
"The Shining | King, Stephen | 1977 | Shelf K1"
```

# --hints--

Unapaswa kutangaza kigezo iitwayo `rawCatalogCards`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+rawCatalogCards\s*=/);
```

`rawCatalogCards` inapaswa kuwa safu ya data.

```js
assert.isArray(rawCatalogCards);
```

Kipengele cha kwanza cha `rawCatalogCards` kinapaswa kuwa `"From a Buick 8 | King, Stephen | 2002 | Shelf K7"`.

```js
assert.strictEqual(rawCatalogCards[0], "From a Buick 8 | King, Stephen | 2002 | Shelf K7");
```

Kipengele cha pili cha `rawCatalogCards` kinapaswa kuwa `"The Shining | King, Stephen | 1977 | Shelf K1"`.

```js
assert.strictEqual(rawCatalogCards[1], "The Shining | King, Stephen | 1977 | Shelf K1");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
