---
id: 699d20623e57e8ffee885454
title: Hatua ya 4
challengeType: 1
dashedName: step-4
---

# --description--

Hatua ya kwanza katika kuchambua kadi ya katalogi ni kutenganisha sehemu nne. Njia ya `.split()` hugawanya mfuatano wa herufi kuwa safu ya data ya vipande vidogo kila inapokutana na kitenzi unachotoa.

Kwa mfano:

```js
"The Shining | King, Stephen | 1977 | Shelf K1".split("|")
// ["The Shining ", " King, Stephen ", " 1977 ", " Shelf K1"]
```

Ndani ya `parseCard`, fanya wito wa `.split("|")` kwa `rawString` na rudisha matokeo. Baada ya hatua hii, konsoli inapaswa kuonyesha safu ya data yenye mfuatano wa herufi nne.

**Kumbuka:** Unaweza kugundua nafasi za ziada karibu na thamani. Utazisafisha katika hatua inayofuata.

# --hints--

Unapaswa kuita `.split("|")` kwa `rawString` ndani ya kitendakazi chako cha `parseCard`.

```js
assert.match(parseCard.toString(), /rawString\s*\.\s*split\s*\(\s*('|"|`)\|\1\s*\)\s*;?/);
```

`parseCard` inapaswa kurudisha safu ya data.

```js
assert.isArray(parseCard("A | B | C | D"));
```

Inapoitwa na `"A | B | C | D"`, `parseCard` inapaswa kurudisha safu ya data yenye vipengele vinne.

```js
assert.strictEqual(parseCard("A | B | C | D").length, 4);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
--fcc-editable-region--
  
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
