---
id: 699d20623e57e8ffee885456
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Sasa baada ya sehemu hizo kuwa safi, unaweza kuzibadilisha kuwa kitu kilicho na muundo. Katika JavaScript, vitu hutumia mabano ya mviringo yenye jozi za "funguo: thamani". Kwa mfano, `{ title: "Dune", year: 1965 }`.

Tangaza vigezo vinne: `title`, `author`, `year`, na `location`. Weka thamani kwa kila moja yao kutoka kipengele kinacholingana cha `trimmedParts`. Kisha sasisha tamko la `return` kurudisha kitu chenye sifa hizo nne.

# --hints--

`parseCard` inapaswa kurudisha kitu chenye sifa za `title`, `author`, `year`, na `location`.

```js
const result = parseCard("The Stand | King, Stephen | 1978 | Shelf K2");
assert.isObject(result);
assert.property(result, "title");
assert.property(result, "author");
assert.property(result, "year");
assert.property(result, "location");
```

Sifa ya `title` ya kitu kinachorudishwa na `parseCard("The Stand | King, Stephen | 1978 | Shelf K2")` inapaswa kuwa na thamani ya `"The Stand"`.

```js
assert.strictEqual(parseCard("The Stand | King, Stephen | 1978 | Shelf K2").title, "The Stand");
```

Sifa ya `author` ya kitu kinachorudishwa na `parseCard("The Stand | King, Stephen | 1978 | Shelf K2")` inapaswa kuwa na thamani ya `"King, Stephen"`.

```js
assert.strictEqual(parseCard("The Stand | King, Stephen | 1978 | Shelf K2").author, "King, Stephen");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
--fcc-editable-region--
  
  return trimmedParts;
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
