---
id: 699d20623e57e8ffee885458
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Sehemu ya mwaka huwasili kama mfuatano wa herufi, lakini unahitaji kuwa nambari kwa ajili ya hesabu inayofuata. Tumia `parseInt()` kuibadilisha.

Mwaka pia unaweza kuwa tupu, kwa hivyo unahitaji njia mbadala. Opereta ya ternary `condition ? valueIfTrue : valueIfFalse` hurudisha moja ya thamani mbili kulingana na sharti.

Sahihisha sifa ya `year` katika taarifa ya return kuwa:

```js
year: year ? parseInt(year) : "Unknown"
```

# --hints--

Unapoitwa na kadi yenye mwaka halali, sifa ya `year` ya kitu kinachorudishwa inapaswa kuwa nambari, si mfuatano wa herufi.

```js
assert.isNumber(parseCard("The Shining | King, Stephen | 1977 | Shelf K1").year);
```

Unapoitwa na `"The Shining | King, Stephen | 1977 | Shelf K1"`, sifa ya `year` inapaswa kuwa `1977`.

```js
assert.strictEqual(parseCard("The Shining | King, Stephen | 1977 | Shelf K1").year, 1977);
```

Unapoitwa na mfuatano wa herufi wenye sehemu ya mwaka tupu, sifa ya `year` inapaswa kuwa `"Unknown"`.

```js
assert.strictEqual(parseCard("X | A |  | Y").year, "Unknown");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
--fcc-editable-region--
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year,
    location: location || "Unknown"
  };
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
