---
id: 699d20623e57e8ffee885455
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Kama unavyoona kutoka kwenye matokeo katika konsoli, baada ya kugawanya, kila sehemu bado ina nafasi za ziada mwanzoni na/au mwishoni. Ili kusafisha kila kipengele, unahitaji kushughulikia kila moja kwa mtu mmoja mmoja.

Hii ni mfano mzuri wa matumizi ya mzunguko wa `for`, ambao hurudia kipande cha msimbo mara moja kwa kila kipengele katika kitu kinachoweza kurudiwa, kama safu ya data:

```js
for (let i = 0; i < array.length; i++) {
  // code to run on each iteration
}
```

`i` huanza kwa `0` (kielezo cha kwanza), mzunguko unaendelea wakati `i` iko chini ya urefu wa safu ya data, na `i++` huongeza `i` kwa moja baada ya kila kurudia.

Ndani ya `parseCard`, tangaza safu ya data tupu iitwayo `trimmedParts`. Kisha ongeza mzunguko wa `for` unaopita juu ya `parts`. Acha mwili wa mzunguko ukiwe tupu kwa sasa, kwani utajaza hatua inayofuata.

# --hints--

Unapaswa kutangaza kigezo iitwayo `trimmedParts` kilichowekwa kuwa safu ya data tupu.

```js
assert.match(parseCard.toString(), /(const|let|var)\s+trimmedParts\s*=\s*\[\s*\]/);
```

Unapaswa kutumia mzunguko wa `for` kurudia juu ya `parts`.

```js
assert.match(parseCard.toString(), /for\s*\(\s*(var|let)\s+(\w+)\s*=\s*0\s*;\s*\2\s*<\s*parts\s*\.\s*length\s*;\s*\2\+\+\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
--fcc-editable-region--
  
--fcc-editable-region--
  return parts;
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
