---
id: 69aa11c2584523e6e991330c
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Njia ya `.trim()` huondoa nafasi tupu kutoka pande zote mbili za mfuatano wa herufi. Kwa mfano, `" hello ".trim()` hurudisha `"hello"`.

Ndani ya mwili wa mzunguko, sukuma `parts[i].trim()` ndani ya `trimmedParts`. Kisha sasisha tamko la `return` lililopo kurudisha `trimmedParts` badala ya `parts`.

Baada ya hatua hii, matokeo yaliyorekodiwa yanapaswa kuonyesha thamani safi bila nafasi za ziada.

# --hints--

Unapaswa kusukuma `parts[i].trim()` ndani ya `trimmedParts` ndani ya mzunguko.

```js
assert.match(parseCard.toString(), /trimmedParts\s*\.\s*push\s*\(\s*parts\s*\[\s*\w+\s*\]\s*\.\s*trim\s*\(\s*\)\s*\)/);
```

`parseCard` inapaswa kurudisha `trimmedParts`.

```js
const explorer = await __helpers.Explorer(code);
const { parseCard } = explorer.allFunctions;
assert.isTrue(parseCard.hasReturn("trimmedParts"));
```

`parseCard` inapaswa kurudisha safu ya data ambapo kila kipengele hakina nafasi tupu mwanzoni wala mwishoni.

```js
assert.isArray(parseCard("  A  |  B  |  C  |  D  "));
```

Inapoitwa na `"  A  |  B  |  C  |  D  "`, kipengele cha kwanza cha safu ya data inayorudishwa kinapaswa kuwa `"A"`.

```js
assert.strictEqual(parseCard("  A  |  B  |  C  |  D  ")[0], "A");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
--fcc-editable-region--
    
  }
  return parts;
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
