---
id: 69757cc0faae0152c1418aad
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在本次研讨会中，你将练习使用数字和数学操作来构建一个账单切分窗口。该工具将在添加餐费和小费后计算每个人应付的金额。

首先，你需要一种方法来跟踪随着成本添加而变化的总金额。在 Python 中，你可以使用一个变量来保存一个随时间变化的整数（整数）。

例如，你可能会写：

```py
my_number = 2
```

创建一个名为 `running_total` 的变量并赋值为 `0`。

# --hints--

你应该有一个名为 `running_total` 的变量。

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('running_total')
    `))
})
```

你应该将整数 `0` 赋值给你的 `running_total` 变量。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('running_total').is_equivalent('running_total = 0')
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
