---
id: 6976cc0b0c135686a4620b94
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

接下来，你需要考虑分摊账单的人数。将此值保存在一个变量中，就像你在上一步中所做的那样。

创建一个名为 `num_of_friends` 的变量，并赋值为 `4`。这将在后续的研讨会中用于计算最终分摊。

# --hints--

你应该有一个名为 `num_of_friends` 的变量。

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('num_of_friends')
    `))
})
```

你应该将整数 `4` 赋值给你的 `num_of_friends` 变量。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('num_of_friends').is_equivalent('num_of_friends = 4')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

--fcc-editable-region--

--fcc-editable-region--
```
