---
id: 6976d8a7cd2a1a5ffcd7f9c4
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

每个课程都有一个费用。你需要将这些金额保存到变量中以便稍后使用。由于这些金额包含分数，你将使用用于表现小数的 float 类型。以下是一个带有 float 值的变量示例：

```py
change = 2.35
```

创建四个变量：`appetizers` 设置为 `37.89`，`main_courses` 设置为 `57.34`，`desserts` 设置为 `39.39`，以及 `drinks` 设置为 `64.21`。 

# --hints--

你应该定义一个名为 `appetizers` 的变量并赋值为 `37.89`。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('appetizers').is_equivalent('appetizers = 37.89')
    `))
})
```

你应该定义一个名为 `main_courses` 的变量，并赋值为 `57.34`。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('main_courses').is_equivalent('main_courses = 57.34')
    `))
})
```

你应该定义一个名为 `desserts` 的变量，并赋值为 `39.39`。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('desserts').is_equivalent('desserts = 39.39')
    `))
})
```

你应该定义一个名为 `drinks` 的变量，并赋值为 `64.21`。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('drinks').is_equivalent('drinks = 64.21')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

--fcc-editable-region--

--fcc-editable-region--
```
