---
id: 6977a4306bb3e856690a4890
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

现在你已经计算了小费，你需要将它添加到你的 `running_total` 中以查找最终账单金额。

使用 `+=` 操作符将 `tip` 的值添加到你的 `running_total`。最后，使用 `print()` 打印字串 `Total with tip:`，后跟一个空格和 `running_total` 的值。

# --hints--

你应该使用 `+=` 操作符将 `tip` 的值添加到你的 `running_total`。

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_stmt('running_total += tip')
    `))
})
```

你应该打印字串 `Total with tip:`，后跟一个空格和 `running_total` 变量。

```js
({
    test: () => runPython(`
import io, contextlib, re

buffer = io.StringIO()
with contextlib.redirect_stdout(buffer):
    exec(_code)

match = re.search(r"Total with tip: ([0-9]+(?:\\.[0-9]+)?)", buffer.getvalue())

assert match
assert abs(float(match.group(1)) - ((37.89 + 57.34 + 39.39 + 64.21) * 1.25)) < 1e-6
    `)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

--fcc-editable-region--

--fcc-editable-region--
```
