---
id: 6977a8bbcbdd557886bd0718
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

现在包含小费后，你就得到了整个团体的最终金额。你必须通过将总账单除以朋友人数来确定每个人欠多少钱。 

在 Python 中，你使用斜线 `/` 来执行除法。例如：

```py
half = 10 / 2
```

创建一个名为 `final_bill` 的变量，并将 `running_total` 除以 `num_of_friends` 的结果赋值给它。 

最后，使用 `print()` 函数显示字串 `Bill per person:`，后跟一个空格和 `final_bill` 的值。

# --hints--

你应该定义一个名为 `final_bill` 的变量。

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('final_bill')
    `))
})
```

你应该使用 `/` 操作符将 `running_total` 除以 `num_of_friends`，并将结果赋值给你的 `final_bill` 变量。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('final_bill').is_equivalent('final_bill = running_total / num_of_friends')
    `))
})
```

你应该打印字串 `Bill per person:`，后跟一个空格和 `final_bill` 变量。

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Bill per person:', final_bill)") or _Node(_code).has_call("print(f'Bill per person: {final_bill}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

running_total += tip
print('Total with tip:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
