---
id: 69779dcb44016eccc05c15a4
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

服务非常好，所以团队决定留下 25％ 的小费。要在 Python 中计算百分比，你可以将总数乘以百分比的小数等价物。

例如，要查找一个值的 10％，你可以使用 `*` 操作符将其乘以 `0.10`：

```py
tax = total * 0.10
```

创建一个名为 `tip` 的变量，并将 `running_total` 乘以 `0.25` 的结果赋值给它。

最后，使用 `print()` 显示字串 `Tip amount:`，后跟一个空格和你的 `tip` 变量的值。

# --hints--

你应该有一个名为 `tip` 的变量。

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('tip')
    `))
})
```

你的 `tip` 变量应是 `running_total * 0.25` 的结果。

```js
({
    test: () => runPython(`
    t = _Node(_code).find_variable('tip')
    assert t.is_equivalent('tip = running_total * 0.25') or t.is_equivalent('tip = 0.25 * running_total')`)
})
```

你应该打印字串 `Tip amount:`，后跟一个空格和 `tip` 变量。

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Tip amount:', tip)") or _Node(_code).has_call("print(f'Tip amount: {tip}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total = appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
