---
id: 6982684f3a25f379e195a5fc
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

现在你已经保存了各项费用，你可以计算总额。

查全率 `+=` 操作符将一个值添加到现有的变量并同时更新它。例如：

```py
total = 10
total += 2 + 2 + 1
print(total)  # total is now 15
```

使用一次 `+=` 操作符，将 `appetizers`、`main_courses`、`desserts` 和 `drinks` 添加到 `running_total`。

最后，使用 `print()` 打印字串 `Total bill so far:`，后跟一个空格和 `running_total` 的值。

**注意：**你可能会注意到输出的小数位数比预期的多。正如你在之前的课程中学到的，这种情况发生是因为数字以二进制形式保存，许多十进制数值无法在这种形式中被精确表现，这导致了四舍五入的误差。

# --hints--

你应该使用 `running_total` 操作符更新 `+=` 一次，以添加所有四门课程的费用。

```js
({
       test: () => runPython(`
import itertools  
perms = itertools.permutations(['+ appetizers', '+ main_courses', '+ desserts', '+ drinks'])  
values = (' '.join(perm).lstrip('+') for perm in perms)  
solutions = (f'running_total += {v}' for v in values)  
assert any(_Node(_code).has_stmt(s) for s in solutions)`)
})
```

你应该打印字串 `Total bill so far:`，后跟一个空格和 `running_total` 的值。

```js
({
    test: () => runPython(`
import io, contextlib, re

buffer = io.StringIO()
with contextlib.redirect_stdout(buffer):
    exec(_code)

match = re.search(r"Total bill so far: ([0-9]+(?:\\.[0-9]+)?)", buffer.getvalue())

assert match
assert abs(float(match.group(1)) - (37.89 + 57.34 + 39.39 + 64.21)) < 1e-6
`)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

--fcc-editable-region--

--fcc-editable-region--
```
