---
id: a6e40f1041b06c996f7b2406
title: بناء دالة للعثور على أول عنصر
challengeType: 26
dashedName: build-a-first-element-finder
---

# --description--

في هذا المختبر، ستنشئ دالة تبحث خلال مصفوفة وتُرجع أول عنصر ينجح في دالة اختبار (المعروفة أيضًا باسم "اختبار الحقيقة").

تُكرر الدالة خلال المصفوفة وتختبر كل عنصر باستخدام دالة الاختبار المقدمة. في النهاية، تُرجع أول عنصر تُعيد دالة الاختبار له القيمة `true`.

مثال:

```js
findElement([1, 3, 5, 8], num => num % 2 === 0) // returns 8
findElement([1, 3, 5], num => num % 2 === 0)    // returns undefined
```

**الهدف:** حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تملك دالة باسم `findElement` تقبل مصفوفة ودالة كمعلمات.
2. يجب أن تُرجع الدالة أول عنصر في المصفوفة ينجح في اختبار الحقيقة. هذا يعني أنه عند استدعاء الدالة الممررة `func` مع عنصر `x`، ينجح الاختبار إذا كانت `func(x)` تُعيد `true`.
3. إذا لم ينجح أي عنصر في الاختبار، يجب أن تُرجع الدالة `undefined`.

# --hints--

يجب أن تملك دالة `findElement`.

```js
assert.isFunction(findElement);
```

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })` يجب أن تُرجع `8`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })` يجب أن تُرجع `undefined`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

`findElement([1, 2, 3, 4], function(num) { return num > 2; })` يجب أن تُرجع `3`.

```js
assert.strictEqual(
  findElement([1, 2, 3, 4], function (num) {
    return num > 2;
  }),
  3
);
```

`findElement(["hello", "world", "javascript"], function(str) { return str.length > 5; })` يجب أن تُرجع `"javascript"`.

```js
assert.strictEqual(
  findElement(["hello", "world", "javascript"], function (str) {
    return str.length > 5;
  }),
  "javascript"
);
```

`findElement(["cat", "dog", "bird"], function(str) { return str.length > 10; })` يجب أن تُرجع `undefined`.

```js
assert.strictEqual(
  findElement(["cat", "dog", "bird"], function (str) {
    return str.length > 10;
  }),
  undefined
);
```

`findElement([2, 4, 6, 8], function(num) { return num % 2 === 0; })` يجب أن تُرجع `2`.

```js
assert.strictEqual(
  findElement([2, 4, 6, 8], function (num) {
    return num % 2 === 0;
  }),
  2
);
```

`findElement([], function(num) { return num > 0; })` يجب أن تُرجع `undefined`.

```js
assert.strictEqual(
  findElement([], function (num) {
    return num > 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
