---
id: 63b61490e633a22b4593e62f
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

按照相同的模式，将 `#add-entry` 元素分配给 `addEntryButton`，将 `#clear` 元素分配给 `clearButton`，并将 `#output` 元素分配给 `output`。

# --hints--

你应该声明一个名为 `addEntryButton` 的变量。

```js
assert.isDefined(addEntryButton);
```

你应该使用 `document.getElementById()` 来获取 `#add-entry` 元素。

```js
assert.match(code, /document\.getElementById\(\s*('|")add-entry\1\s*\)/g);
```

你应该将 `#add-entry` 元素分配给 `addEntryButton`。

```js
assert.deepEqual(addEntryButton, document.getElementById('add-entry'));
```

你应该声明一个名为 `clearButton` 的变量。

```js
assert.isDefined(clearButton);
```

你应该使用 `document.getElementById()` 来获取 `#clear` 元素。

```js
assert.match(code, /document\.getElementById\(\s*('|")clear\1\s*\)/g);
```

你应该将 `#clear` 元素分配给 `clearButton`。

```js
assert.deepEqual(clearButton, document.getElementById('clear'));
```

你应该声明一个名为 `output` 的变量。

```js
// .isDefined cannot be used for `output`, as browser defines such variable on it own.
assert.match(code, /const\s+output/);
```

你应该使用 `document.getElementById()` 来获取 `#output` 元素。

```js
assert.match(code, /document\.getElementById\(\s*('|")output\1\s*\)/g);
```

你应该将 `#output` 元素分配给 `output`。

```js
// assert.deepEqual(output, document.getElementById('output')); cannot be used,
// as browser defines `output` variable on it own for `#output`.
assert.match(code, /const\s+output\s*=\s*document.getElementById\(('|")output\1\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
          <div class="controls">
            <span>
              <label for="entry-dropdown">Add food or exercise:</label>
              <select id="entry-dropdown" name="options">
                <option value="breakfast" selected>Breakfast</option>
                <option value="lunch">Lunch</option>
                <option value="dinner">Dinner</option>
                <option value="snacks">Snacks</option>
                <option value="exercise">Exercise</option>
              </select>
              <button type="button" id="add-entry">Add Entry</button>
            </span>
          </div>
          <div>
            <button type="submit">
              Calculate Remaining Calories
            </button>
            <button type="button" id="clear">Clear</button>
          </div>
        </form>
        <div id="output" class="output hide"></div>
      </div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js
--fcc-editable-region--
const calorieCounter = document.getElementById('calorie-counter');
const budgetNumberInput = document.getElementById('budget');
const entryDropdown = document.getElementById('entry-dropdown');

--fcc-editable-region--
```
