---
id: 66ea7adae8053065a64f9002
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop, you will review working with the different `Math` object methods by building a Mathbot. This Mathbot will log some math operations and messages to the console in efforts to teach you about JavaScript's `Math` object.

For this first step, start by creating a variable called `botName` and assign it the string value of `"MathBot"`.

Then, create another variable called `greeting` and assign it the sentence `"Hi there! My name is [botName goes here] and I am here to teach you about the Math object!"`. In place of the `[botName goes here]` placeholder, use the `botName` variable.

You are free to use template literals or string concatenation with the `+` operator to achieve this.

Finally, log the `greeting` variable to the console.

# --hints--

Ви повинні мати змінну під назвою `botName`.

```js
assert.isNotNull(botName);
```

Змінна `botName` має бути рядком.

```js
assert.isString(botName);
```

Призначте значення рядка `"MathBot"` до змінної `botName`.

```js
assert.strictEqual(botName, "MathBot");
```

Ви повинні мати змінну під назвою `greeting`.

```js
assert.isNotNull(greeting);
```

Змінна `greeting` має бути рядком.

```js
assert.isString(greeting);
```

You should assign the following sentence to the `greeting` variable: `"Hi there! My name is [MathBot goes here] and I am here to teach you about the Math object!"`. Make sure to replace `[MathBot goes here]` with the `botName` variable.

```js
assert.strictEqual(greeting, "Hi there! My name is MathBot and I am here to teach you about the Math object!");
```

You should log the `greeting` variable to the console.

```js
assert.match(code, /console\.log\(greeting\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
