---
id: 68c2fb153760d075321c5da3
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Camperbot ameunda ukurasa wa blogu wa safari ya uandishi wa programu, lakini inaonekana ukurasa una matatizo ya ufikikaji na mbinu mbaya. Kazi yako katika warsha hii ni kurekebisha matatizo haya kwa Camperbot.

Kundi la kwanza la makosa linahusiana na matumizi ya vichwa vya sehemu. Katika funzo la awali, ulijifunza kuwa ni mbinu bora kutumia `h1` moja tu kwa kila ukurasa. Lakini inaonekana Camperbot anatumia `h1` kadhaa.

Acha `<h1>Welcome to Camperbot's Blog</h1>` kama ilivyo na badilisha vipengele vingine vya `h1` kuwa vipengele vya `h2`.

# --hints--

Haupaswi kubadilisha `<h1>Welcome to Camperbot's Blog</h1>`.

```js
const heading = document.querySelectorAll("h1")[0];
assert.equal(heading?.textContent, "Welcome to Camperbot's Blog");
```

Unapaswa kubadilisha kipengele cha `h1` chenye maandishi ya `My Journey Learning to Code` kuwa kipengele cha `h2`.

```js
const firstPost = document.getElementById("post1");
assert.equal(firstPost?.tagName, "H2");
assert.equal(firstPost?.textContent, "My Journey Learning to Code");
```

Unapaswa kubadilisha kipengele cha `h1` chenye maandishi ya `Accessibility Matters` kuwa kipengele cha `h2`.

```js
const secondPost = document.getElementById("post2");
assert.equal(secondPost?.tagName, "H2");
assert.equal(secondPost?.textContent, "Accessibility Matters");
```

Unapaswa kubadilisha kipengele cha `h1` chenye maandishi ya `What's Next?` kuwa kipengele cha `h2`.

```js
const thirdPost = document.getElementById("post3");
assert.equal(thirdPost?.tagName, "H2");
assert.equal(thirdPost?.textContent, "What's Next?");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Camperbot's Blog</title>
</head>
<body>
  <h1>Welcome to Camperbot's Blog</h1>

  <h2>Navigation</h2>
  <ul>
    <li><a href="#post1">My Journey</a></li>
    <li><a href="#post2">Accessibility</a></li>
    <li><a href="#post3">Next Steps</a></li>
  </ul>
  <div>
    --fcc-editable-region--
    <h1 id="post1">My Journey Learning to Code</h1>
    <p>I started learning to code a few months ago and it's been a wild ride!</p>
    
    <h4>Early Challenges</h4>
    <p>At first, syntax was really confusing.</p>
    
    <h4>Breakthroughs</h4>
    <p>Eventually things started to click.</p>
  </div>

  <div>
    <h1 id="post2">Accessibility Matters</h1>
    <p>Today I learned that not everyone uses the web the same way I do.</p>

    <h5>Screen Readers</h5>
    <p>These tools help visually impaired users browse websites.</p>
  </div>

  <div>
    <h1 id="post3">What's Next?</h1>
    --fcc-editable-region--
    <p>I'm excited to dive into JavaScript and build interactive features!</p>

    <h3>Coming soon: My first JavaScript project!</h3>
    <p>Stay tuned for some exciting interactive blog features.</p>
  </div>
  <h3>Contact Me</h3>
  <p>Email me at camperbot@blog.io</p>
</body>
</html>
```
