---
id: 5d5a813321b9e3db6c106a46
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

JavaScript 是一種強大的語言，它允許你構建交互式網站。

*注意*：對於本課程中的所有剩餘項目，你需要具備基本的 HTML 和 CSS 知識。 如果你不熟悉 HTML 和 CSS，請完成[響應式網頁設計認證](https: www.freecodecamp.org/learn/2022/responsive-web-design/)。

首先，使用 `DOCTYPE`、`html`、`head` 和 `body` 標籤創建你的標準 HTML 模版，然後爲 `meta` 添加一個 `charset` 屬性。 添加 `title` 元素，將它的文本設置爲 `RPG - Dragon Repeller`。 添加一個 `link` 標籤，鏈接到 `styles.css` 文件，以添加樣式表。

最後，在你的 `body` 中創建一個 `div` 元素，將其 `id` 設置爲 `game`。

# --hints--

你的代碼應該包含 `DOCTYPE` 引用。

```js
assert(code.match(/<!DOCTYPE/gi));
```

你應在 `DOCTYPE` 後添加一個空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你應將文檔類型定義爲 `html` 。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

在類型後，你應該用 `>` 來關閉 `DOCTYPE` 聲明。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你的 `html` 元素應該有一個開始標籤。 不要忘記設置 `lang` 屬性。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素應該有一個結束標籤。

```js
assert(code.match(/<\/html\s*>/));
```

你的 `DOCTYPE` 聲明應位於 HTML 的開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你應該有一個 `head` 開始標籤。

```js
assert(code.match(/<head\s*>/i));
```

你應該有一個 `head` 結束標籤。

```js
assert(code.match(/<\/head\s*>/i));
```

你應該有一個 `body` 開始標籤。

```js
assert(code.match(/<body\s*>/i));
```

你應該有一個 `body` 結束標籤。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素應該是同級元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

你的代碼應該有一個 `meta` 元素。

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

你的 `meta` 元素應該有一個 `charset` 屬性，其值爲 `UTF-8`。

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

你的代碼應該有一個 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

您應該有一個結尾的 `title` 標記。

```js
assert(code.match(/<\/title\s*>/i));
```

你的代碼中應該有一個 `link` 元素。

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

你應該創建一個 `div` 元素。

```js
const div = document.querySelector('div');
assert.exists(div);
```

你的 `div` 元素應該有一個值爲 `game` 的 `id`。

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

你的 `div` 元素應該在 `body` 元素內。

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
