---
id: 62a115879a6d51422652cbfc
title: 步驟11
challengeType: 0
dashedName: step-11
---

# --description--

在你的角色扮演遊戲中，用戶將能夠跟蹤他們的統計數據，購買武器並與怪物戰鬥。 在繼續遊戲的交互式 JavaScript 部分之前，你需要先創建將顯示遊戲信息的 HTML 元素。

在 `#game` 元素中創建四個 `div` 元素。 給它們分配各自的 `id` 值，依次爲：`stats`、`controls`、`monsterStats`和`text`。

# --hints--

應該創建四個新的 `div` 元素。

```js
assert.equal(document.querySelectorAll('div')?.length, 5);
```

你應該給新的 `div` 元素之一一個 `id` 爲 `stats`。

```js
assert.exists(document.querySelector('div#stats'));
```

你應該給新的 `div` 元素之一一個 `id` 爲 `controls`。

```js
assert.exists(document.querySelector('div#controls'));
```

你應該給新的 `div` 元素之一一個 `id` 爲 `monsterStats`。

```js
assert.exists(document.querySelector('div#monsterStats'));
```

你應該給新的 `div` 元素之一一個 `id` 爲 `text`。

```js
assert.exists(document.querySelector('div#text'));
```

應該以正確的順序放置新的 `div` 元素。

```js
function __t(a, b) {
  return document.querySelector(a)?.nextElementSibling?.getAttribute('id') === b;
}
assert(__t('div#stats','controls') && __t('div#controls','monsterStats') && __t('div#monsterStats','text'));
```

應該將新的 `div` 元素放在 `#game` 元素中。

```js
assert.equal(document.querySelector('#game')?.children?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
