---
id: 62a257659d0d1e2456f24ba2
title: 步驟4
challengeType: 0
dashedName: step-4
---

# --description--

在開始編寫項目代碼之前，你應該將其移動到其自己的文件中以保持井然有序。

刪除你的 `console.log("Hello World");` 行。 然後將現在爲空的 `script` 元素的 `src` 屬性設置爲 `./script.js`。

# --hints--

你的代碼中不應該有 `console.log("Hello World");` 行。

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

你的 `script` 元素應將 `src` 屬性設置爲 `./script.js`。

```js
const script = document.querySelector("script[data-src$='script.js']");
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>
      console.log("Hello World");
    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
