---
id: 62a23c1d505bfa13747c8a9b
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

將數字 `0`、`100` 和 `50` 包含在 `span` 元素中，並將這些新的 `span` 元素包含在 `strong` 元素中。 然後分別爲新的 `span` 元素設置 `id` 值爲 `xpText`、`healthText` 和 `goldText` 。

你的答案應該遵循這個基本結構：

```html
<span class="stat">TEXT <strong><span id="VALUE">TEXT</span></strong></span>
```

# --hints--

你應該在第一個 `.stat` 元素中添加一個 `strong` 元素。

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

你的第一個新的 `strong` 元素應該有一個 `span` 元素。

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

你的第一個嵌套 `span` 元素應該具有 `xpText` 的 `id`。

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'xpText');
```

你的第一個 `span` 元素應該包裹在文本 `0` 周圍。

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '0');
```

你的第一個 `.stat` 元素的文本仍應爲 `XP: 0`。

```js
const stat = document.querySelectorAll('.stat')[0];
assert.equal(stat.innerText.trimEnd(), 'XP: 0');
```

你應該在第二個 `.stat` 元素中添加一個 `strong` 元素。

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

你的第二個新 `strong` 元素應該有一個 `span` 元素。

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

你的第二個嵌套 `span` 元素應該具有 `healthText` 的 `id`。

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'healthText');
```

你的第二個 `span` 元素應包裹文本 `100`。

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '100');
```

你的第二個 `.stat` 元素的文本仍應爲 `Health: 100`。

```js
const stat = document.querySelectorAll('.stat')[1];
assert.equal(stat.innerText.trimEnd(), 'Health: 100');
```

你應該在第三個 `.stat` 元素中添加一個 `strong` 元素。

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

你的第三個新 `strong` 元素應該有一個 `span` 元素。

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

你的第三個嵌套 `span` 元素應該具有 `goldText` 的 `id`。

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'goldText');
```

你的第三個 `span` 元素應包裹文本 `50`。

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '50');
```

你的第三個 `.stat` 元素的文本仍應爲 `Gold: 50`。

```js
const stat = document.querySelectorAll('.stat')[2];
assert.equal(stat.innerText.trimEnd(), 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
--fcc-editable-region--
        <span class="stat">XP: 0</span>
        <span class="stat">Health: 100</span>
        <span class="stat">Gold: 50</span>
--fcc-editable-region--
      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
