---
id: 62a3a7e4f1060e2fc5ffb34b
title: 步驟 8
challengeType: 0
dashedName: step-8
---

# --description--

創建另一個名爲 `currentWeaponIndex` 的變量並將其設置爲 `0`。

# --hints--

你應該用關鍵字 `let` 來聲明一個名字叫做 `currentWeaponIndex` 的變量。

```js
assert.match(code, /let\s+currentWeaponIndex/i);
```

你應使用駝峯命名法來命名你的變量。

```js
assert.match(code, /currentWeaponIndex/);
```

你的 `currentWeaponIndex` 變量值應設爲 `0`。

```js
assert.equal(currentWeaponIndex, 0);
```

你應將你的變量的值初始化爲 `0`。

```js
assert.match(code, /let\s+currentWeaponIndex\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
