---
id: 62a3a0a3c0a4b32915d26a6e
title: 步驟6
challengeType: 0
dashedName: step-6
---

# --description--

刪除 `console.log("Hello World");` ，然後開始寫你自己的項目代碼。

在上一個項目中，你學習瞭如何使用變量和 `let` 關鍵字，如下所示：

```js
let camperbot;
```

你還學會了如何初始化像這樣值的變量：

```js
let age = 32;
```

使用 `let` 關鍵字聲明一個名爲 `xp` 的變量，併爲其分配數字 `0` 的值。

# --hints--

你的代碼中不應該有 `console.log("Hello World");` 行。

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

你應該使用 `let` 關鍵字聲明變量。

```js
assert.match(code, /let/);
```

你應該聲明一個名爲 `xp` 的變量。

```js
assert.match(code, /let\s+xp/);
```

`xp` 變量值應爲 `0`。

```js
assert.equal(xp, 0);
```

你應該將 `xp` 變量初始化爲 `0`。

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
