---
id: 62a3b41c9494f937560640ab
title: 步驟 10
challengeType: 0
dashedName: step-10
---

# --description--

再聲明兩個名爲 `monsterHealth` 和 `inventory` 的變量。

對於你的 `inventory` 變量，爲其分配包含字符串 `"stick"` 的數組的值。

想想你在之前的項目中學會了像這樣數組方法：

```js
let exampleArray = ["first", "second", "third"];
```

# --hints--

你應該使用 `let` 關鍵詞聲明你的 `monsterHealth` 變量。

```js
assert.match(code, /let\s+monsterHealth/i);
```

你應該使用駝峯命名法（camelCase）來命名 `monsterHealth` 變量。

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` 不應該有值。

```js
assert.isUndefined(monsterHealth);
```

你不應該爲你的 `monsterHealth` 變量賦值。

```js
assert.match(code, /let\s+monsterHealth\s*;?/);
```

你應該使用 `let` 關鍵詞聲明你的 `inventory` 變量。

```js
assert.match(code, /let\s+inventory/i);
```

你的 `inventory` 變量仍然應該是一個數組。

```js
assert.isArray(inventory);
```

你的 `inventory` 變量只有1個值。

```js
assert.lengthOf(inventory, 1);
```

你的 `inventory` 變量應存放着 `"stick"` 字符串。

```js
assert.include(inventory, "stick");
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
--fcc-editable-region--
```
