---
id: 62a8d81f539f004776dd9b1e
title: 步驟 110
challengeType: 0
dashedName: step-110
---

# --description--

現在你可以啓動代碼來對抗怪物了。 爲了使你的代碼保持井然有序，你的 `fightDragon` 函數已被移動，以便靠近其他 `fight` 函數。

在 `weapons` 數組下方，定義一個 `monsters` 變量併爲其分配一個數組。 將該數組設置爲具有三個對象，每個對象均具有 `name`、`level` 和 `health` 屬性。 第一個對象的值應該按順序爲 `"slime"`、`2` 和 `15`。 第二個應該是 `"fanged beast"`、`8` 和 `60`。 第三個應該是 `"dragon"`、`20` 和 `300`。

# --hints--

你應該使用 `const` 來聲明 `monsters` 變量。

```js
assert.match(code, /const\s+monsters/);
```

你的 `monsters` 變量應該是一個數組。

```js
assert.isArray(monsters);
```

你的 `monsters` 變量應該有 3 個值。

```js
assert.lengthOf(monsters, 3);
```

你的 `monsters` 數組應該有 3 個對象。

```js
assert(monsters.every(val => typeof val === "object"));
```

`monsters` 數組中的第一個值應該是一個對象，其 `name` 屬性設置爲 `"slime"`。

```js
assert.equal(monsters[0].name, "slime");
```

`monsters` 數組中的第一個值應該是一個對象，其 `level` 屬性設置爲 `2`。

```js
assert.equal(monsters[0].level, 2);
```

`monsters` 數組中的第一個值應該是一個對象，其 `health` 屬性設置爲 `15`。

```js
assert.equal(monsters[0].health, 15);
```

`monsters` 數組中的第二個值應該是一個對象，其 `name` 屬性設置爲 `"fanged beast"`。

```js
assert.equal(monsters[1].name, "fanged beast");
```

`monsters` 數組中的第二個值應該是一個對象，其 `level` 屬性設置爲 `8`。

```js
assert.equal(monsters[1].level, 8);
```

`monsters` 數組中的第二個值應該是一個對象，其 `health` 屬性設置爲 `60`。

```js
assert.equal(monsters[1].health, 60);
```

`monsters` 數組中的第三個值應該是一個對象，其 `name` 屬性設置爲 `"dragon"`。

```js
assert.equal(monsters[2].name, "dragon");
```

`monsters` 數組中的第三個值應該是一個對象，其 `level` 屬性設置爲 `20`。

```js
assert.equal(monsters[2].level, 20);
```

`monsters` 數組中的第三個值應該是一個對象，其 `health` 屬性設置爲 `300`。

```js
assert.equal(monsters[2].health, 300);
```

你不應該向你的對象添加任何額外的屬性。

```js
assert.deepEqual(monsters, [
  {name: "slime", level: 2, health: 15},
  {name: "fanged beast", level: 8, health: 60},
  {name: "dragon", level: 20, health: 300}
])
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");
const weapons = [
  { name: 'stick', power: 5 },
  { name: 'dagger', power: 30 },
  { name: 'claw hammer', power: 50 },
  { name: 'sword', power: 100 }
];
--fcc-editable-region--

--fcc-editable-region--
const locations = [
  {
    name: "town square",
    "button text": ["Go to store", "Go to cave", "Fight dragon"],
    "button functions": [goStore, goCave, fightDragon],
    text: "You are in the town square. You see a sign that says \"Store\"."
  },
  {
    name: "store",
    "button text": ["Buy 10 health (10 gold)", "Buy weapon (30 gold)", "Go to town square"],
    "button functions": [buyHealth, buyWeapon, goTown],
    text: "You enter the store."
  },
  {
    name: "cave",
    "button text": ["Fight slime", "Fight fanged beast", "Go to town square"],
    "button functions": [fightSlime, fightBeast, goTown],
    text: "You enter the cave. You see some monsters."
  }
];

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

function update(location) {
  button1.innerText = location["button text"][0];
  button2.innerText = location["button text"][1];
  button3.innerText = location["button text"][2];
  button1.onclick = location["button functions"][0];
  button2.onclick = location["button functions"][1];
  button3.onclick = location["button functions"][2];
  text.innerText = location.text;
}

function goTown() {
  update(locations[0]);
}

function goStore() {
  update(locations[1]);
}

function goCave() {
  update(locations[2]);
}

function buyHealth() {
  if (gold >= 10) {
    gold -= 10;
    health += 10;
    goldText.innerText = gold;
    healthText.innerText = health;
  } else {
    text.innerText = "You do not have enough gold to buy health.";
  }
}

function buyWeapon() {
  if (currentWeaponIndex < weapons.length - 1) {
    if (gold >= 30) {
      gold -= 30;
      currentWeaponIndex++;
      goldText.innerText = gold;
      let newWeapon = weapons[currentWeaponIndex].name;
      text.innerText = "You now have a " + newWeapon + ".";
      inventory.push(newWeapon);
      text.innerText += " In your inventory you have: " + inventory;
    } else {
      text.innerText = "You do not have enough gold to buy a weapon.";
    }
  } else {
    text.innerText = "You already have the most powerful weapon!";
    button2.innerText = "Sell weapon for 15 gold";
    button2.onclick = sellWeapon;
  }
}

function sellWeapon() {
  if (inventory.length > 1) {
    gold += 15;
    goldText.innerText = gold;
    let currentWeapon = inventory.shift();
    text.innerText = "You sold a " + currentWeapon + ".";
    text.innerText += " In your inventory you have: " + inventory;
  } else {
    text.innerText = "Don't sell your only weapon!";
  }
}

function fightSlime() {

}

function fightBeast() {

}

function fightDragon() {
  console.log("Fighting dragon.");
}
```
