---
id: 62fc20387ef88d1d1998aac5
title: 步驟 18
challengeType: 0
dashedName: step-18
---

# --description--

使用 `querySelector()` 來獲取另外兩個 `button` 元素，並使用它們的 `id`：`button2` 和 `button3`。 將它們存儲在名爲 `button2` 和 `button3` 的變量中。 請記住使用 `const`。

# --hints--

你應該使用 `const` 聲明一個 `button2` 變量。

```js
assert.match(code, /const\s+button2/);
```

你的 `button2` 變量應該具有 `"#button2"` 元素的值。

```js
assert.deepEqual(button2, document.querySelector('#button2'));
```

你應該使用 `const` 聲明一個 `button3` 變量。

```js
assert.match(code, /const\s+button3/);
```

你的 `button3` 變量應該具有 `"#button3"` 元素的值。

```js
assert.deepEqual(button3, document.querySelector('#button3'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--
const button1 = document.querySelector("#button1");

--fcc-editable-region--
```
