---
id: 6894c1b715b8d18a18ef3203
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Neste workshop, você terá prática com as classes e objetos do Python construindo um inventário de instrumentos musicais.

Para definir uma classe em Python, use a palavra-chave `class` seguida do nome da classe e dois pontos.

```python
class ClassName:
    # class contents go here
```

Lembre-se de que em Python, os nomes de classes normalmente usam PascalCase, o que significa que cada palavra começa com uma letra maiúscula sem espaços entre as palavras.

Comece criando uma classe chamada `MusicalInstrument`. Adicione `pass` como o corpo da classe por enquanto.

# --hints--

Você deve criar uma classe chamada `MusicalInstrument`.

```js
({
 test: () => assert(runPython(`_Node(_code).has_class("MusicalInstrument")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
