---
id: 6895e3f91c069e0068a42582
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

Dentro do seu método `play`, use `print()` para exibir uma mensagem sobre tocar o instrumento.

Use uma f-string para incluir o nome do instrumento na mensagem: `The [instrument name] is fun to play!`. O `[instrument name]` deve ser substituído pelo nome real do instrumento.

# --hints--

Seu método `play` deve usar uma f-string para exibir a mensagem `The [instrument name] is fun to play!`. Use `self.name` para acessar o nome do instrumento.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_call('print(f"The {self.name} is fun to play!")')`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        pass
--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
