---
id: 6895e572df8596062234f571
title: Passo 13
challengeType: 20
dashedName: step-13
---

# --description--

Finalmente, faça o mesmo para `instrument_2`. Chame o método `play()` em `instrument_2` e depois imprima o resultado de chamar `get_fact()` em `instrument_2`.

Isto completará seu workshop de instrumentos musicais!

# --hints--

Você deve chamar o método `play()` na instância `instrument_2`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_stmt("instrument_2.play()")`))
});
```

Você deve chamar `get_fact()` em `instrument_2` e imprimir o resultado.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.get_fact())")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


--fcc-editable-region--
instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

instrument_1.play()
print(instrument_1.get_fact())

--fcc-editable-region--
```

# --solutions--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')


instrument_1.play()              # The Oboe is fun to play!
print(instrument_1.get_fact())   # The Oboe is part of the woodwind family of instruments.

instrument_2.play()              # The Trumpet is fun to play!
print(instrument_2.get_fact())   # The Trumpet is part of the brass family of instruments.
```
