---
id: 6895e4940e39950288e406c2
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

Em seguida, adicione outro método à sua classe chamado `get_fact` que retorna um fato sobre o instrumento.

Este método deve ter apenas o parâmetro `self` e retornar uma f-string que diz `The [instrument name] is part of the [instrument type] family of instruments.`

Ao contrário do método `play` que imprime uma mensagem diretamente, este método deve retornar uma string que pode ser usada em outra parte do seu código.

# --hints--

Você deve definir um método chamado `get_fact` com um parâmetro `self`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").has_args("self")`))
});
```

Seu método `get_fact` deve retornar uma f-string contendo o nome e o tipo do instrumento.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").find_return().is_equivalent("return f'The {self.name} is part of the {self.instrument_type} family of instruments.'")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
