---
id: 6895e5337b1914050b0af7bd
title: Passo 12
challengeType: 20
dashedName: step-12
---

# --description--

Agora, chame o método `get_fact()` na instância `instrument_1` e imprima o resultado.

Como `get_fact()` retorna uma string em vez de imprimi-la diretamente, você precisa usar `print()` para exibir a string retornada.

# --hints--

Você deve chamar `get_fact()` em `instrument_1` e imprimir o resultado.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.get_fact())")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


--fcc-editable-region--
instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

instrument_1.play()

--fcc-editable-region--
```
