---
id: 6895e4e023e07603bcef46f9
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Agora vamos usar os métodos que você criou. 

Chame o método `play()` na sua instância `instrument_1`. Isso irá imprimir a mensagem sobre tocar o Oboe.

Para chamar um método em uma instância, use a notação de ponto:

```python
instance_name.method_name()
```

# --hints--

Você deve chamar o método `play()` na instância `instrument_1`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_stmt("instrument_1.play()")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


--fcc-editable-region--
instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
```
