---
id: 6895e356fd877efd6e3657ee
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--


Dentro da classe `MusicalInstrument`, defina outro método chamado `play` com apenas o parâmetro `self`. O parâmetro `self` permite que os métodos acessem os atributos da instância específica.

# --hints--

Você deve definir um método chamado `play` com um parâmetro `self`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_args("self")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
