---
id: 6895dffc64e97df0e62e9ef9
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Métodos são funções definidas dentro de uma classe que podem executar ações usando os dados da classe. Aqui está como você cria métodos:

```python
def method_name(parameter1, parameter2,...):
    # Method contents will go here
```

Em Python, o método `__init__` é um método especial conhecido como método inicializador, que é chamado quando você cria uma nova instância da classe.

Dentro da sua classe `MusicalInstrument`, adicione um método `__init__` com três parâmetros: `self`, `name` e `instrument_type`.

O parâmetro `self` é uma forma padrão de se referir à instância da classe e é obrigatório como o primeiro parâmetro para todos os métodos de instância em uma classe.

# --hints--

Você deve definir um método `__init__` dentro da sua classe `MusicalInstrument`.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").has_function("__init__")`))
});
```

Seu método `__init__` deve ter três parâmetros: `self`, `name` e `instrument_type`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").has_args("self, name, instrument_type")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    pass
--fcc-editable-region--
```
