---
id: 6895e0bdb0348bf46952ab67
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Agora é hora de criar instâncias da sua classe `MusicalInstrument`.

Para criar uma instância de uma classe, você chama o nome da classe como uma função, passando os argumentos necessários para o método `__init__` (exceto `self`, que é tratado automaticamente pelo Python).

```python
my_instance = ClassName(arg1, arg2)
```

Quando você cria essa instância, o Python chama seu método `__init__` e armazena os valores como atributos da instância que pertencem apenas a essa instância específica.

Fora da definição da classe, crie uma instância chamada `instrument_1` para um `Oboe` que é um instrumento `woodwind`.

# --hints--

Você deve criar uma instância da classe `MusicalInstrument` chamada `instrument_1`.

```js
({
    test: () => runPython(`
    assert isinstance(instrument_1, MusicalInstrument)
    `)
})
```

Você deve passar `Oboe` e `woodwind` como argumentos ao criar `instrument_1`.

```js
({ 
    test: () => runPython(`
        assert instrument_1.name == "Oboe"
        assert instrument_1.instrument_type == "woodwind" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--
```
