---
id: 6895e197180c9af6988e1e19
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Crie uma segunda instância chamada `instrument_2` para um `Trumpet` que é um instrumento `brass`.

Observe como você pode criar múltiplas instâncias a partir do mesmo template de classe. Cada instância armazena seus próprios valores de atributos de forma independente das outras instâncias.

# --hints--

Você deve criar uma instância chamada `instrument_2` da classe `MusicalInstrument`.

```js
({
    test: () => runPython(`
    assert isinstance(instrument_2, MusicalInstrument)
    `)
})
```

Você deve passar `Trumpet` e `brass` como argumentos ao criar `instrument_2`.

```js
({ 
    test: () => runPython(`
        assert instrument_2.name == "Trumpet"
        assert instrument_2.instrument_type == "brass" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')

--fcc-editable-region--
```
