---
id: 6895e0452b3370f26f365a46
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

No método `__init__`, você precisa atribuir os parâmetros aos atributos da instância. Atributos da instância são variáveis que pertencem a uma instância específica de uma classe.

Dentro do seu método `__init__`, atribua o parâmetro `name` a `self.name` e o parâmetro `instrument_type` a `self.instrument_type`.

# --hints--

Você deve atribuir o parâmetro `name` a `self.name` dentro do método `__init__`.

```js
 ({
    test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[0].is_equivalent("self.name = name")`))
});
```

Você deve atribuir o parâmetro `instrument_type` a `self.instrument_type` dentro do método `__init__`.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[1].is_equivalent("self.instrument_type = instrument_type")`))
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        pass
--fcc-editable-region--
```
