---
id: 6895e235935106f88281b648
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Você pode acessar os atributos das suas instâncias usando a notação de ponto:

```py
instance_name.attribute_name
```

Isso permite que você recupere os valores que foram armazenados durante a inicialização.

Tente imprimir os atributos `name` e `instrument_type` de ambas as suas instâncias. Você poderá ver os valores que atribuiu quando as criou.

# --hints--

Você deve imprimir o atributo `name` de `instrument_1`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

Você deve imprimir o atributo `instrument_type` de `instrument_1`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

Você deve imprimir o atributo `name` de `instrument_2`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

Você deve imprimir o atributo `instrument_type` de `instrument_2`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
```
