---
id: 68c16924498728d259e4781c
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

现在创建一个 `Employee` 类的实例，传入字串 `Charlie Brown` 和 `trainee`。将该实例赋值给名为 `charlie_brown` 的变量。

# --hints--

你应该有一个名为 `charlie_brown` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("charlie_brown")`)) })
```

你应该将 `Employee('Charlie Brown', 'trainee')` 赋值给 `charlie_brown`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("charlie_brown").is_equivalent("charlie_brown = Employee('Charlie Brown', 'trainee')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level


--fcc-editable-region--
```
