---
id: 68c18b1d0fce081f85312d97
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

在之前的课程中，你已经了解到以单个下划线开头的属性按照惯例是供内部使用的。

将 `name` 和 `level` 两个属性修改为 `_name` 和 `_level`，因为这些不应该从它们的类外部被修改。

请注意，这并不阻止在类外部访问或修改该属性。此外，在 Python 中，总有办法访问以双下划线开头的私有属性。

# --hints--

你的 `__init__` 方法应将 `self._name` 设置为 `name`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._name = name")`)) })
```

你的 `__init__` 方法应将 `self._level` 设置为 `level`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level

charlie_brown = Employee('Charlie Brown', 'trainee')
--fcc-editable-region--
```
