---
id: 68c15a2a1f0e70ca7154aba5
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在本次研讨会中，你将构建一个用于员工的薪资跟踪系统。

首先创建一个名为 `Employee` 的类。在其中创建一个带有 `self`、`name` 和 `level` 参数的 `__init__` 方法。在 `__init__` 方法内，将 `name` 和 `level` 赋值给同名的实例属性。

# --hints--

你应该有一个名为 `Employee` 的类。

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Employee")`)) })
```

你的 `Employee` 类应有一个带有三个参数 `self`、`name` 和 `level` 的 `__init__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_args("self, name, level")`)) })
```

你的 `__init__` 方法应将 `self.name` 设置为 `name`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

你的 `__init__` 方法应将 `self.level` 设置为 `level`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
