---
id: 68c28c6517e7de5e85f8ebe5
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

既然你已经为 `_name` 和 `_level` 属性提供了 getter，更新 `__str__` 返回的字串以使用 `self.name` 和 `self.level`。这将调用 getter 而不是直接访问属性。

# --hints--

你的 `__str__` 方法应返回 `f'{self.name}: {self.level}'`

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_return("f'{self.name}: {self.level}'")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self._name}: {self._level}'
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)
print(charlie_brown.level)
```
