---
id: 68c28d659d345270d4a61084
title: 步骤11
challengeType: 20
dashedName: step-11
---

# --description--

`__repr__` 方法是一个特殊的方法，应该返回一个可以用来实例化该对象的字符串表现。

例如，`Employee('Charlie', 'developer')` 的 `__repr__` 方法应返回字串 `Employee('Charlie', 'developer')`，该字串与用于创建该对象的字串相同。

为你的 `Employee` 类添加一个带有 `self` 参数的 `__repr__` 方法，并使其返回一个可用于实例化该对象的字串。

# --hints--

你的 `Employee` 类应有一个带有 `self` 参数的 `__repr__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__repr__").has_args("self")`)) })
```

你的 `__repr__` 方法应返回用于实例化对象的字串。

```js
({ test: () => runPython(`
  empl = Employee("Frank", "dreamer")
  assert repr(empl) == 'Employee("Frank", "dreamer")' or repr(empl) == "Employee('Frank', 'dreamer')"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
