---
id: 68c27fa1697ceb7c578b2c9f
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

`@property` 装饰器在 Python 中用于将方法转换为属性。它通常用于定义 getter 方法，即用于获取属性值的方法：

```py
class Person:
    def __init__(self, name):
        self._name = name

    @property
    def name(self):
        return self._name

p = Person('Alice')
print(p.name)  # Alice
```

创建一个名为 `name` 的方法，带有一个 `self` 参数，并用 `@property` 装饰它。在方法内部，返回 `self._name`。
    
# --hints--

你的 `Employee` 类应有一个带有 `self` 参数的 `name` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_args("self")`)) })
```

你的 `name` 方法应使用 `@property` 装饰。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_decorators("property")`)) })
```

你的 `name` 方法应返回 `self._name`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_return("self._name")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
