---
id: 68c27eb292bfec6b9402d58f
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

为你的 `Employee` 类添加一个 `__str__` 方法。使它返回一个格式为 `name: level` 的 f-字串，将 `name` 和 `level` 替换为相应的属性。

之后，将 `charlie_brown` 打印到控制台。

# --hints--

你的 `Employee` 类应有一个带有 `self` 参数的 `__str__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_args("self")`)) })
```

你的 `__str__` 方法应返回一个格式为 `name: level` 的 f-string，分别用 `self._name` 和 `self._level` 的值替换 `name` 和 `level`。

```js
({ test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    empl = Employee("Frank", "dreamer")
    print(empl)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "Frank: dreamer" in output
`) })
```

你应该将 `charlie_brown` 打印到控制台。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    

charlie_brown = Employee('Charlie Brown', 'trainee')

--fcc-editable-region--
```
