---
id: 68c28b7cc3ab8f4daf144b5e
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

按照你之前的步骤，为 `_level` 属性创建一个 getter。

# --hints--

你的 `Employee` 类应有带有 `self` 参数的 `level` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_args("self")`)) })
```

你的 `level` 方法应使用 `@property` 装饰。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_decorators("property")`)) })
```

你的 `level` 方法应返回 `self._level`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_return("self._level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    
--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)

```
