---
id: 68c281cecda94ba2c6838e17
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

既然你为 `name` 定义了 getter，你可以通过 `name` 属性访问 `_name` 属性。所以将 `charlie_brown.name` 打印到控制台。 

# --hints--

你应该将 `charlie_brown.name` 打印到控制台。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown.name)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    @property
    def name(self):
        return self._name

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
