---
id: 68c2d31a86d5672ca49d4521
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

现在是为 `__init__` 方法添加一些验证的时候了。在该方法的开头，创建一个 `if` 语句，查看 `name` 或 `level` 是否不是 `str` 的实例。

在 `if` 语句中，引起一个带有消息 `'name' and 'level' attribute must be of type 'str'.` 的 `TypeError`。

# --hints--

你应该在你的 `__init__` 方法中有一个 `if` 语句。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0]`)) })
```

你的 `if` 语句应查看 `name` 和 `level` 是否任一不是 `str` 的实例。

```js
({ test: () => runPython(`
  cond = _Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0].find_conditions()[0]
  conditions = [
    "not (isinstance(name, str) and isinstance(level, str))",
    "not isinstance(name, str) or not isinstance(level, str)",
    "not (isinstance(level, str) and isinstance(name, str))",
    "not isinstance(level, str) or not isinstance(name, str)",
  ]
  assert any(cond.is_equivalent(c) for c in conditions)
`) })
```

你的 `if` 语句应引起一个带有消息 `'name' and 'level' attribute must be of type 'str'.` 的 `TypeError`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError("\\'name\\' and \\'level\\' attribute must be of type \\'str\\'.")')`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
--fcc-editable-region--
    def __init__(self, name, level):
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
