---
id: 68c7f7012a700243eff1cbc0
title: 步骤 16
challengeType: 20
dashedName: step-16
---

# --description--

用于实例化员工的等级应从特定等级中选择。你将使用 `_base_salaries` 来验证等级并为员工设置正确的薪资。

在你现有的 `if` 语句之后，创建另一个 `if` 来查看 `level` 是否不在 `Employee._base_salaries` 中。

在新的 `if` 语句中，引起一个带有消息 `Invalid value '{level}' for 'level' attribute.` 的 `ValueError`，其中 `{level}` 应替换为 `level` 参数的值。

# --hints--

当 `level` 不在 `Employee._base_salaries` 中时，你应该引起一个带有消息 `Invalid value '{level}' for 'level' attribute.` 的 `ValueError`，其中 `{level}` 应替换为 `level` 参数的值。

```js
({ test: () => runPython(`
  try:
    Employee("Frank", "dreamer")
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
    assert False, "Expected to raise ValueError with invalid level"
`) })
```

当 `level` 在 `Employee._base_salaries` 中时，你不应引起任何异常。

```js
({ test: () => runPython(`
  levels = [
    "trainee",
    "junior",
    "mid-level",
    "senior"
  ]
  for level in levels:
    try:
      Employee("Frank", level)
    except Exception:
      assert False, f"Expected not to raise ValueError with level '{level}'"
    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }
--fcc-editable-region--
    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
