---
id: 68c8002beb157fe47d2e7699
title: 步骤 20
challengeType: 20
dashedName: step-20
---

# --description--

setter 是用于设置属性值的方法，允许进行验证检查和限制。你可以使用 `@propertyName.setter` 装饰器来创建 setter，其中 `propertyName` 应该匹配要设置的属性名称：

```py
class Person:
    def __init__(self, name):
        self._name = name

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, value):
        self._name = value

p = Person('Alice')
p.name = 'Abigail' # Calls the setter
print(p.name) # Abigail
```

在你的 getter 方法之后，创建一个带有参数 `self` 和 `new_name` 的 `name` 方法。用 `@name.setter` 装饰该方法。在方法内部，将 `self._name` 设置为 `new_name`。

# --hints--

你的 `Employee` 类应有一个带有参数 `self` 和 `new_name` 的 `name` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_args("self, new_name")`)) })
```

你的 `name` 方法应使用 `@name.setter` 装饰。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_decorators("name.setter")`)) })
```

你的 `name` 方法应将 `self._name` 设置为 `new_name`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_stmt("self._name = new_name")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    
--fcc-editable-region--
    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary


charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
