---
id: 68c809c990f253912a9e9209
title: 步骤21
challengeType: 20
dashedName: step-21
---

# --description--

正如你在之前的课程中学到的，setter 提供了一种控制属性如何被修改的方法。为了确保 `new_name` 是正确的类型，创建一个 `if` 语句，当 `new_name` 不是 `str` 的实例时，引起一个带有消息 `'name' must be a string.` 的 `TypeError`。

# --hints--

你应该在你的 `name` setter 中有一个 `if` 语句。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].find_ifs()[0]`)) })
```

当 `new_name` 不是 `str` 的实例时，你应该引起一个带有消息 `'name' must be a string.` 的 `TypeError`。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  non_names = [0, True, []]
  for i in non_names:
    try:
      emp.name = i
    except TypeError as e:
      assert str(e) == "'name' must be a string."
    else:
      assert False, "Expected to raise TypeError with non-string new_name"
`) })
```

当 `new_name` 是字串时，你不应引起任何异常。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.name = "Jack"
  except Exception:
    assert False, "Expected not to raise any exception with valid new_name"
    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name
--fcc-editable-region--
    @name.setter
    def name(self, new_name):
        
        self._name = new_name
--fcc-editable-region--
    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
