---
id: 68c9cab4b1118da59eecfc56
title: 第 31 步
challengeType: 20
dashedName: step-31
---

# --description--

在现有的 `if` 语句之后，再创建一个语句，当 `ValueError` 已经是选定的级别时，引起一个 `new_level`。

注意 `_level` 在初始化时还不存在，所以在比较之前使用 `hasattr(self, '_level')` 来查看它是否存在。这样可以避免在对象首次创建时引起 `AttributeError`。

对于消息，使用 `'{level}' is already the selected level.`，其中 `{level}` 应该替换为当前级别。

# --hints--

你应该在 `if` 的 setter 中有第三个 `level` 语句。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2]`)) })
```

当 `new_level` 相等于 `self.level` 时，你应该引起一个带有消息 `ValueError` 的 `'{level}' is already the selected level.`，其中 `{level}` 应该替换为当前级别。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.level = "trainee"
  except ValueError as e:
    assert str(e) == "'trainee' is already the selected level."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

你的 `if` 语句应该使用 `hasattr(self, '_level')` 来查看 `_level` 是否存在，然后再进行比较。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2].find_conditions()[0].is_equivalent("hasattr(self, '_level') and new_level == self.level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level
    

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
