---
id: 68c9c3aa714bc326e026b826
title: 步骤 26
challengeType: 20
dashedName: step-26
---

# --description--

在设置新级别之前，必须查看它是否为有效级别。在你的 setter 开始处，创建一个 `if` 语句，当 `new_level` 不是 `Employee._base_salaries` 的密钥时，引起 `ValueError`。

对于错误信息，使用 `Invalid value '{new_level}' for 'level' attribute.`，其中 `{new_level}` 应替换为传递给 setter 的参数。

# --hints--

你应该在你的 `level` setter 中有一个 `if` 语句。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[0]`)) })
```

当 `new_level` 不是 `Employee._base_salaries` 的密钥时，你应该引起一个带有消息 `Invalid value '{new_level}' for 'level' attribute.` 的 `ValueError`，其中 `{new_level}` 应替换为传递给 setter 的参数。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.level = "dreamer"
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

当 `new_level` 是 `Employee._base_salaries` 的密钥时，你不应引起任何异常

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  for i in Employee._base_salaries:
    try:
      emp.level = i
    except Exception:
      assert False, "Expected not to raise any exception with valid new_level"    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level
--fcc-editable-region--
    @level.setter
    def level(self, new_level):
        
        self._level = new_level
    
--fcc-editable-region--
    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
