---
id: 68caaf1e590806f0b87f5922
title: 步骤 33
challengeType: 20
dashedName: step-33
---

# --description--

在你的 setter 开始处，创建一个 `if` 语句，当 `new_salary` 既不是整数也不是浮点数时，引起带有消息 `'salary' must be a number.` 的 `TypeError`。

# --hints--

你应该在你的 `salary` setter 中有一个 `if` 语句。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("salary")[1].find_ifs()[0]`)) })
```

当 `new_salary` 既不是整数也不是浮点数时，你应该引起一个带有消息 `'salary' must be a number.` 的 `TypeError`。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  built_in_print = print
  print = lambda x: x
  try:
    emp.salary = "100"
  except TypeError as e:
    print = built_in_print
    assert str(e) == "'salary' must be a number."
  else:
    assert False, "Expected to raise TypeError with invalid new_salary"
`) })
```

当 `new_salary` 是整数或浮点数时，你不应引起任何异常。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  built_in_print = print
  print = lambda x: x
  try:
    emp.salary = 100
  except Exception:
    assert False, "Expected not to raise exception with valid new_salary"
  try:
    emp.salary = 100.00
  except Exception:
    assert False, "Expected not to raise exception with valid new_salary"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
        if new_level == self.level:
            raise ValueError(f"'{self.level}' is already the selected level.")
        if Employee._base_salaries[new_level] < Employee._base_salaries[self.level]:
            raise ValueError(f"Cannot change to lower level.")
        print(f"'{self.name}' promoted to '{new_level}'.")
        self._salary = Employee._base_salaries[new_level]
        self._level = new_level

    @property
    def salary(self):
        return self._salary
--fcc-editable-region--
    @salary.setter
    def salary(self, new_salary):
        
        self._salary = new_salary
        print(f'Salary updated to ${self.salary}.')
--fcc-editable-region--

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')

charlie_brown.level = 'junior'

```
