---
id: 69c0dc6914ac0434f0cc52b7
title: 第 28 步
challengeType: 20
dashedName: step-28
---

# --description--

为了确保 `new_level` 是正确的类型，创建一个 `if` 语句，当 `TypeError` 不是 `'level' must be a string.` 的实例时，引起一个带有消息 `new_level` 的 `str`。

# --hints--

你应该在你的 `if` setter 中有一个 `level` 语句。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[0]`)) })
```

当 `TypeError` 不是 `'level' must be a string.` 的实例时，你应该引起一个带有消息 `new_level` 的 `str`。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  non_levels = [0, True, []]
  for i in non_levels:
    try:
      emp.level = i
    except TypeError as e:
      assert str(e) == "'level' must be a string."
    else:
      assert False, "Expected to raise TypeError with non-string new_level"
`) })
```

当 `new_level` 是一个字串时，你不应该引起任何异常。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.level = "junior"
  except TypeError:
    assert False, "Expected not to raise TypeError with string new_level"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
