---
id: 66cc172af7159a7c67804544
title: ステップ 12
challengeType: 1
dashedName: step-12
---

# --description--

これで計算機に掛け算の機能を追加する時です。

`calculateProduct` という名前の関数を宣言し、2つのパラメータ `num1` と `num2` を受け取るようにしてください。

関数の中で、2つのパラメータの積を返してください。

その後、`console.log` を追加し、`calculateProduct` 関数を引数 `13` と `5` で呼び出してください。

# --hints--

`calculateProduct` という名前の関数があるはずです。

```js
assert.isFunction(calculateProduct);
```

`calculateProduct` 関数は `num1` と `num2` という2つのパラメータを持つべきです。

```js
const regex = __helpers.functionRegex('calculateProduct', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

`calculateProduct` 関数は2つのパラメータの積を返すべきです。

```js
assert.equal(calculateProduct(5, 2), 10);
assert.equal(calculateProduct(10, 5), 50);
```

`console.log` が `calculateProduct` 関数を引数 `13` と `5` で呼び出しているはずです。

```js
assert.match(code, /console\.log\s*\(\s*calculateProduct\s*\(\s*13\s*,\s*5\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

--fcc-editable-region--

--fcc-editable-region--
```
