---
id: 66cc0ba5881acb7692cfc4de
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

次の関数を追加する時間です。

`addThreeAndFour` という名前の関数を宣言し、`3` と `4` の合計を返してください。

その後、`addThreeAndFour` の中で `console.log` 関数を呼び出して結果を確認してください。

# --hints--

`addThreeAndFour` という名前の関数を宣言してください。

```js
assert.isFunction(addThreeAndFour);
```

`addThreeAndFour` 関数は `3` と `4` の合計を返す必要があります。

```js
assert.strictEqual(addThreeAndFour(), 7);
```

`addThreeAndFour` 関数は `+` 演算子を使って合計を返す必要があります。

```js
const explorer = await __helpers.Explorer(code);
const { addThreeAndFour } = explorer.allFunctions;
const validExpressions = ["3 + 4", "4 + 3", "(3 + 4)", "(4 + 3)"];
assert.isTrue(validExpressions.some(expr => addThreeAndFour?.hasReturn(expr)));
```

`addThreeAndFour` 関数を `console.log` の中で呼び出してください。

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(addThreeAndFour())";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
