---
id: 66cc1a3a39aef47d6473cb2f
title: ステップ 13
challengeType: 1
dashedName: step-13
---

# --description--

次のステップは、計算機に割り算の機能を追加することです。

`calculateQuotient` という名前の関数を宣言し、2つのパラメータ `num1` と `num2` を受け取ってください。

関数の中で、2つのパラメータの商を返してください。

その後、`console.log` を追加し、`calculateQuotient` 関数を引数 `7` と `11` で呼び出してください。

# --hints--

`calculateQuotient` という名前の関数があるはずです。

```js
assert.isFunction(calculateQuotient);
```

`calculateQuotient` 関数は `num1` と `num2` という2つのパラメータを持つべきです。

```js
const regex = __helpers.functionRegex('calculateQuotient', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

`calculateQuotient` 関数は2つのパラメータの商を返すべきです。

```js
assert.equal(calculateQuotient(10, 2), 5);
assert.equal(calculateQuotient(100, 10), 10);
```

`console.log` が `calculateQuotient` 関数を引数 `7` と `11` で呼び出しているはずです。

```js
assert.match(code, /console\.log\s*\(\s*calculateQuotient\s*\(\s*7\s*,\s*11\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

--fcc-editable-region--

--fcc-editable-region--
```
