---
id: 66cc0ba5881acb7692cfc4de
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

次に、別の関数を追加する時間です。

`addThreeAndFour` という名前の関数を宣言し、`3` と `4` の合計を返してください。

その後、`addThreeAndFour` の中で `console.log` 関数を呼び出して結果を確認してください。

# --hints--

`addThreeAndFour` という名前の関数を宣言してください。

```js
assert.isFunction(addThreeAndFour);
```

`addThreeAndFour` 関数は `3` と `4` の合計を返す必要があります。

```js
assert.equal(addThreeAndFour(), 7);
```

`addThreeAndFour` 関数は `+` 演算子を使って合計を返す必要があります。

```js
const declarationStyle = /return\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const arrowStyle = /=>\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const combinedRegex = new RegExp(`(?:${declarationStyle.source})|(?:${arrowStyle.source})`);
assert.match(code, combinedRegex);
```

`addThreeAndFour` 関数を `console.log` の中で呼び出してください。

```js
assert.match(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
