---
id: 66cc14ca14e65e7a9ebecb80
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

計算機が2つの数字を足せるようになったので、次は2つの数字を引き算する関数を作りましょう。

`calculateDifference` という名前の関数を宣言し、2つのパラメータ `num1` と `num2` を受け取るようにしてください。

関数の中で、2つのパラメータの差を返してください。

# --hints--

`calculateDifference` という名前の関数があるはずです。

```js
assert.isFunction(calculateDifference);
```

`calculateDifference` 関数は `num1` と `num2` という2つのパラメータを持つべきです。

```js
const regex = __helpers.functionRegex('calculateDifference', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

`calculateDifference` 関数は2つのパラメータの差を返すべきです。

```js
assert.equal(calculateDifference(5, 2), 3);
assert.equal(calculateDifference(10, 5), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

--fcc-editable-region--

--fcc-editable-region--
```
