---
id: 66cc10958d72e57907f898b2
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

`calculateSum` 関数の中で、2つのパラメータの合計を返す必要があります。

# --hints--

`calculateSum` 関数は2つのパラメータの合計を返すべきです。

```js
assert.equal(calculateSum(1, 2), 3);
assert.equal(calculateSum(5, 5), 10);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
