---
id: 66cc21d23238dc8240a8a182
title: ステップ16
challengeType: 1
dashedName: step-16
---

# --description--

計算機に数の二乗を計算してほしいですね。数の二乗とは、その数を自分自身で掛けたものです。

JavaScriptで数の二乗を計算するには、<dfn>Math.pow()</dfn>メソッドを使えます。`Math.pow()`メソッドは2つの引数を取ります：底の数と指数です。

また、指数演算子（`**`）を使って数の二乗を計算することもできます。

例を示します：

```js
// base number is 5 
// the exponent is 2
Math.pow(5, 2); // 25

// using the exponentiation operator
5 ** 2; // 25
```

`calculateSquare`パラメータを取り、`num`の二乗を返す`num`という関数を宣言してください。

# --hints--

`calculateSquare`という関数を用意してください。

```js
assert.isFunction(calculateSquare);
```

`calculateSquare`関数は`num`パラメータを取るべきです。

```js
const regex = __helpers.functionRegex('calculateSquare', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

`calculateSquare`関数は`num`の二乗を返すべきです。

```js
assert.equal(calculateSquare(5), 25);
assert.equal(calculateSquare(10), 100);
assert.equal(calculateSquare(3), 9);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

--fcc-editable-region--

--fcc-editable-region--
```
