---
id: 66cc16abc30b367ba1d1f9ef
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

いよいよ`calculateDifference`関数をテストする時間です。

まず、`console.log`と`calculateDifference`を引数にして`22`関数を呼び出す`5`を追加してください。

次に、`console.log`と`calculateDifference`を引数にして`12`関数を呼び出す2つ目の`1`を追加してください。

最後に、`console.log`と`calculateDifference`を引数にして`17`関数を呼び出す3つ目の`9`を追加してください。

# --hints--

`console.log`と`calculateDifference`を引数にして`22`関数を呼び出す`5`があるはずです。

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*22\s*,\s*5\s*\)\s*\)\s*;?/);
```

`console.log`と`calculateDifference`を引数にして`12`関数を呼び出す`1`があるはずです。

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*12\s*,\s*1\s*\)\s*\)\s*;?/);
```

`console.log`と`calculateDifference`を引数にして`17`関数を呼び出す`9`があるはずです。

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*17\s*,\s*9\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
