---
id: 66cc0a9e06e00b75d6782be9
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

前のレッスンでは、関数を呼び出す方法を学びました。関数を呼び出すとは、その関数の中のコードを実行することを意味します。

関数を呼び出す方法をもう一度確認しましょう：

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

`console.log` 文を追加し、その中で `addTwoAndSeven` 関数を呼び出してください。

これで `2` と `7` の合計がコンソールに表示されるはずです。

# --hints--

コードの中に `console.log` があるはずです。

```js
assert.match(code, /console\.log(.*)/);
```

`addTwoAndSeven` の中で `console.log` 関数を呼び出す必要があります。

```js
assert.match(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}
--fcc-editable-region--

--fcc-editable-region--
```
