---
id: 66cc0de2920bec775f610424
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

関数は期待通りに動作していますが、コードに繰り返しが多くあります。

計算機を作っているので、足し合わせたい数字のあらゆる組み合わせごとに関数を作る必要はありません。

どんな2つの数字でも足し合わせられる、再利用可能な単一の関数を作るほうが良いです。

これまで書いたコードはすべて削除してください。

次のステップでは、関数のパラメータと引数の扱い方をレビューします。

# --hints--

`addTwoAndSeven` という名前の関数を作ってはいけません。

```js
const explorer = await __helpers.Explorer(code);
const { addTwoAndSeven } = explorer.allFunctions;
assert.notExists(addTwoAndSeven);
```

`addTwoAndSeven` という関数の呼び出しをしてはいけません。

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const unexpectedStatement = "console.log(addTwoAndSeven())";
assert.notInclude(strippedCode, unexpectedStatement);
```

`addThreeAndFour` という名前の関数を作ってはいけません。

```js
const explorer = await __helpers.Explorer(code);
const { addThreeAndFour } = explorer.allFunctions;
assert.notExists(addThreeAndFour);
```

`addThreeAndFour` という関数の呼び出しをしてはいけません。

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const unexpectedStatement = "console.log(addThreeAndFour())";
assert.notInclude(strippedCode, unexpectedStatement);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
