---
id: 66cc12fa504b0479dac479a0
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

関数の利点は、異なる引数で呼び出せるため、同じコードをさまざまな値で再利用できることです。

以下は、異なる文字列引数で関数を呼び出す例です。

```js
// function definition
function greetUser(name) {
  console.log(`Hello ${name}!`);
}

// function calls
greetUser("John"); // "Hello John!"
greetUser("Jane"); // "Hello Jane!"
```

`console.log` を追加して、引数 `calculateSum` と `2` で `5` 関数を呼び出してください。

# --hints--

`console.log` が引数 `calculateSum` と `2` で `5` 関数を呼び出しているはずです。

```js
assert.match(code, /console\.log\(calculateSum\(\s*2\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
