---
id: 66cc095dfe1682753d2ab030
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

前のレッスンで学んだ概念の一つに、関数から戻り値を返す方法があります。

関数から戻り値を返す方法をもう一度確認しましょう：

```js
function myFunction() {
  return "Hello World";
}
```

あなたの`addTwoAndSeven`関数の中で、`2`と`7`の合計を返してください。

# --hints--

あなたの`addTwoAndSeven`関数は、`2`と`7`の合計を返す必要があります。

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

あなたの`addTwoAndSeven`関数は、`+`演算子を使って合計を返す必要があります。

```js
assert.match(code, /return\s*(2\s*\+\s*7|7\s*\+\s*2)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {  
--fcc-editable-region--

--fcc-editable-region--
}
```
