---
id: 66cc0f1ae40802781b2ea972
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

前のレッスンでは、関数のパラメータの扱い方を学びました。

関数のパラメータとは、関数の宣言で定義され、プレースホルダーとして機能する変数のことです。

以下はパラメータを持つ関数の例です：

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

`calculateSum` という名前の関数を宣言し、`num1` と `num2` の2つのパラメータを受け取るようにしてください。

# --hints--

`calculateSum` という名前の関数を宣言してください。

```js
assert.isFunction(calculateSum);
```

`calculateSum` 関数は `num1` と `num2` の2つのパラメータを受け取る必要があります。

```js
const explorer = await __helpers.Explorer(code);
const { calculateSum } = explorer.allFunctions;
assert.lengthOf(calculateSum?.parameters, 2);
assert.equal(calculateSum?.parameters[0].toString(), "num1");
assert.equal(calculateSum?.parameters[1].toString(), "num2");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
