---
id: 66cc26907ff6908402af0149
title: ステップ 17
challengeType: 1
dashedName: step-17
---

# --description--

`calculateSquare` 関数をテストするには、その関数を数回呼び出す必要があります。

`console.log` を追加して、引数に `calculateSquare` を渡して `2` 関数を呼び出してください。

次に、引数に `console.log` を渡して `calculateSquare` 関数を呼び出す別の `9` を追加してください。

# --hints--

引数に `console.log` を渡して `calculateSquare` 関数を呼び出す `2` があるはずです。

```js
assert.match(code, /console\.log\s*\(\s*calculateSquare\s*\(\s*2\s*\)\s*\)/);
```

引数に `console.log` を渡して `calculateSquare` 関数を呼び出す `9` があるはずです。

```js
assert.match(code, /console\.log\s*\(\s*calculateSquare\s*\(\s*9\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

--fcc-editable-region--

--fcc-editable-region--
```
