---
id: 66cc06012e46aa6bc9b8c001
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、計算機アプリを作りながら関数の使い方をレビューします。

これまでのレッスンで、次のように関数を宣言する方法を学びました。

```js
// regular function
function myFunction() {

}

// arrow function
const myArrowFunction = () => {

}
```

まず、`addTwoAndSeven` という名前の関数を作成してください。

通常の関数構文でもアロー関数構文でもどちらでも選べます。

# --hints--

`addTwoAndSeven` という名前の関数を宣言してください。

```js
assert.isFunction(addTwoAndSeven);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
