---
id: 66cc2739f687e484d50bb6f1
title: ステップ 18
challengeType: 1
dashedName: step-18
---

# --description--

電卓の最後の機能は平方根の機能です。

平方とは、ある数を自分自身で掛けた数のことです。したがって、平方根とは、自分自身で掛けると元の数になる数のことです。

例えば、`9` の平方根は `3` です。なぜなら `3 * 3 = 9` だからです。

JavaScriptで数の平方根を求めるには、<dfn>Math.sqrt()</dfn> メソッドを使います。

例を示します。

```js
// result: 3
Math.sqrt(9);
```

`calculateSquareRoot` という名前の関数を宣言し、`num` というパラメータを受け取り、`num` の平方根を返してください。

# --hints--

`calculateSquareRoot` という名前の関数を作成してください。

```js
assert.isFunction(calculateSquareRoot);
```

`calculateSquareRoot` 関数は `num` というパラメータを受け取る必要があります。

```js
const regex = __helpers.functionRegex('calculateSquareRoot', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

`calculateSquareRoot` 関数は `num` の平方根を返す必要があります。

```js
assert.equal(calculateSquareRoot(9), 3);
assert.equal(calculateSquareRoot(16), 4);
assert.equal(calculateSquareRoot(25), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

--fcc-editable-region--

--fcc-editable-region--
```
