---
id: 66cdf76685e4cb5a8726e27b
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

`calculateSum` 関数をさらにテストするために、異なる引数で呼び出す必要があります。

`console.log` をもう一つ追加して、`calculateSum` 関数を引数 `10` と `10` で呼び出してください。

その `console.log` の下に、`console.log` をもう一つ追加して、`calculateSum` 関数を引数 `5` と `5` で呼び出してください。

# --hints--

`console.log` が `calculateSum` 関数を引数 `10` と `10` で呼び出しているはずです。

```js
assert.match(code, /console\.log\(calculateSum\(10\s*,\s*10\)\)/);
```

`console.log` が `calculateSum` 関数を引数 `5` と `5` で呼び出しているはずです。

```js
assert.match(code, /console\.log\(calculateSum\(5\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

console.log(calculateSum(2, 5));

--fcc-editable-region--

--fcc-editable-region--
```
