---
id: 587d825c367417b2b2512c8f
title: マージソートの実装
challengeType: 1
forumTopicId: 301614
dashedName: implement-merge-sort
---

# --description--

もう一つのよく使われる中級のソートアルゴリズムはマージソートです。クイックソートと同様に、マージソートも分割統治法の再帰的な方法で配列をソートします。マージソートは、2つの配列がそれぞれソートされていれば、それらをソートするのは比較的簡単であるという事実を利用します。しかし、最初は1つの配列だけが入力として与えられるので、どうやって2つのソート済み配列に分けるのでしょうか？それは、元の入力を再帰的に半分に分割し、要素が1つだけの配列になるまで続けることで実現します。要素が1つの配列は自然にソート済みなので、そこから結合を始められます。この結合が元の配列を分割した再帰呼び出しを巻き戻し、最終的にすべての要素がソートされた配列を作り出します。マージソートのステップは次の通りです：

**1)** 入力配列を再帰的に半分に分割し、要素が1つだけのサブ配列を作る。

**2)** ソート済みのサブ配列をそれぞれマージして、最終的なソート済み配列を作る。

マージソートは効率的なソート方法で、時間計算量は *O(nlog(n))* です。このアルゴリズムは性能が良く、比較的実装しやすいため人気があります。

ちなみに、ここで扱うソートアルゴリズムはこれが最後です。ただし、後のツリー型データ構造のセクションで、二分ヒープを使って実装するもう一つの効率的なソート方法であるヒープソートについて説明します。

**指示:** 整数の配列を入力として受け取り、最小から最大の順にソートされた整数の配列を返す関数`mergeSort`を書いてください。実装の良い方法は、例えば2つのソート済み配列をマージする役割の`merge`という関数と、マージに渡すために要素が1つの配列を作る再帰処理を担当する`mergeSort`という関数の2つを書くことです。頑張ってください！

# --hints--

`mergeSort`は関数であるべきです。

```js
assert.isFunction(mergeSort);
```

`mergeSort`はソート済みの配列（最小から最大）を返すべきです。

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    mergeSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`mergeSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])`は順序以外は変更されていない配列を返すべきです。

```js
assert.sameMembers(
  mergeSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`mergeSort`は組み込みの`.sort()`メソッドを使わないでください。

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    mergeSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function mergeSort(array) {
  if (array.length === 1) {
    return array;
  } else {
    const splitIndex = Math.floor(array.length / 2);
    return merge(
      mergeSort(array.slice(0, splitIndex)),
      mergeSort(array.slice(splitIndex))
    );
  }

  // Merge two sorted arrays
  function merge(array1, array2) {
    let merged = [];
    while (array1.length && array2.length) {
      if (array1[0] < array2[0]) {
        merged.push(array1.shift());
      } else if (array1[0] > array2[0]) {
        merged.push(array2.shift());
      } else {
        merged.push(array1.shift(), array2.shift());
      }
    }

    // After looping ends, one array is empty, and other array contains only
    // values greater than all values in `merged`
    return [...merged, ...array1, ...array2];
  }
}

mergeSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]);
```
