---
id: 587d825a367417b2b2512c89
title: クイックソートの実装
challengeType: 1
forumTopicId: 301615
dashedName: implement-quick-sort
---

# --description--

ここでは中級のソートアルゴリズムであるクイックソートに進みます。クイックソートは効率的な再帰的分割統治法で、配列をソートします。この方法では、元の配列からピボット値を選びます。配列はピボット値より小さい値と大きい値の2つのサブ配列に分割されます。次に、両方のサブ配列に対してクイックソートのアルゴリズムを再帰的に呼び出した結果を結合します。これを空の配列または要素が1つの配列になるまで繰り返し、その配列を返します。再帰呼び出しが戻ることでソート済みの配列が得られます。

クイックソートは非常に効率的なソート方法で、平均して*O(nlog(n))*の性能を発揮します。また比較的簡単に実装できます。これらの特徴から、人気があり有用なソート方法です。

**指示:** `quickSort` という関数を作成してください。この関数は整数の配列を入力として受け取り、最小から最大までの順にソートされた整数の配列を返します。ピボット値の選び方は重要ですが、ここではどのピボットでも構いません。簡単のため、最初または最後の要素を使うことができます。

# --hints--

`quickSort` は関数である必要があります。

```js
assert(typeof quickSort == 'function');
```

`quickSort` はソートされた配列（最小から最大）を返す必要があります。

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    quickSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`quickSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` は順序以外は変更されていない配列を返す必要があります。

```js
assert.sameMembers(
  quickSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quickSort` は組み込みの `.sort()` メソッドを使わないでください。

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quickSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function quickSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function quickSort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quickSort(lesser), ...equal, ...quickSort(greater)];
  }
}
```
