---
id: 587d8259367417b2b2512c86
title: 挿入ソートを実装する
challengeType: 1
forumTopicId: 301613
dashedName: implement-insertion-sort
---

# --description--

次に見るソート方法は挿入ソートです。この方法はリストの先頭にソート済みの配列を作りながら進めます。最初の要素でソート済み配列を始めます。次に次の要素を調べて、ソート済み配列の中で正しい位置になるまで後ろに向かって入れ替えます。リストの最後までこの操作を繰り返し、新しい要素をソート済み部分に後ろ向きに入れ替え続けます。このアルゴリズムは平均および最悪の場合で二次の時間計算量を持ちます。

**指示:** `insertionSort` という関数を作成してください。この関数は整数の配列を入力として受け取り、最小から最大までの順にソートされた整数の配列を返します。

# --hints--

`insertionSort` は関数である必要があります。

```js
assert.isFunction(insertionSort);
```

`insertionSort` はソートされた配列（最小から最大）を返す必要があります。

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    insertionSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`insertionSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` は順序以外は変更されていない配列を返す必要があります。

```js
assert.sameMembers(
  insertionSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`insertionSort([5, 4, 33, 2, 8])` は `[2, 4, 5, 8, 33]` を返す必要があります。

```js
assert.deepEqual(insertionSort([5, 4, 33, 2, 8]), [2, 4, 5, 8, 33])
```

`insertionSort` は組み込みの `.sort()` メソッドを使わないでください。

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    insertionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function insertionSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function insertionSort (array) {
  for (let currentIndex = 0; currentIndex < array.length; currentIndex++) {
    let current = array[currentIndex];
    let j = currentIndex - 1;
    while (j > -1 && array[j] > current) {
      array[j + 1] = array[j];
      j--;
    }
    array[j + 1] = current;
  }
  return array;
}
```
