---
id: 61abc7ebf3029b56226de5b6
title: バイナリサーチを実装する
challengeType: 1
forumTopicId: 487618
dashedName: implement-binary-search
---

# --description--

バイナリサーチは、ソート済み配列から要素を検索するための**O(log(n))**の効率を持つアルゴリズムです。以下のステップで動作します：

1. ソート済み配列の中央の`value`を見つけます。もし`value == target`なら`true`を返します（値が見つかり検索が完了したことを意味します）。
1. 中央の`value < target`なら、次の比較では配列の右半分を検索します。
1. 中央の`value > target`なら、次の比較では配列の左半分を検索します。
1. 配列全体を検索しても値が存在しない場合は、`false`を返します（配列を検索し値が存在しないことを意味します）。

ご覧の通り、配列を繰り返し半分にしていくため、log(n)の効率が得られます。このチャレンジでは、どのようにターゲット値にたどり着いたか、つまりたどった経路を示してください！

# --instructions--

バイナリサーチアルゴリズムを配列に対して実装し、ターゲットを見つけるまでにたどった経路（各中央値の比較）を返す関数`binarySearch`を書いてください。

この関数は、ソート済みの整数配列とターゲット値を入力として受け取ります。元の配列を半分に分割するたびに見つけた中央の値を順番に含む配列を返します。ターゲット値は返される配列の最後の要素であるべきです。値が見つからなければ、文字列`Value Not Found`を返してください。

例えば、`binarySearch([1,2,3,4,5,6,7], 5)`は`[4,6,5]`を返します。

このチャレンジでは、半分に分割するときに必ず`Math.floor()`を使って割り算してください：`Math.floor(x/2)`。これにより一貫したテスト可能な経路が得られます。

**注意：** テストには以下の配列が使われます：

```js
const testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
```

# --hints--

`binarySearch`は関数であるべきです。

```js
assert.isFunction(binarySearch);
```

`binarySearch(testArray, 0)`は`[13, 5, 2, 0]`を返すべきです。

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 0), [13, 5, 2, 0]);
```

`binarySearch(testArray, 1)`は`[13, 5, 2, 0, 1]`を返すべきです。

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 1), [13, 5, 2, 0, 1]);
```


`binarySearch(testArray, 2)`は`[13, 5, 2]`を返すべきです。

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 2), [13, 5, 2]);
```

`binarySearch(testArray, 6)`は文字列`Value Not Found`を返すべきです。

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.strictEqual(binarySearch(_testArray, 6), 'Value Not Found');
```

`binarySearch(testArray, 11)`は`[13, 5, 10, 11]`を返すべきです。

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 11), [13, 5, 10, 11])
```

`binarySearch(testArray, 13)`は`[13]`を返すべきです。

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 13), [13]);
```

`binarySearch(testArray, 70)`は`[13, 19, 22, 49, 70]`を返すべきです。

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 70), [13, 19, 22, 49, 70]);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let arrayPath = [];
  return arrayPath;
}
```



# --solutions--

```js
let binarySearch = (searchList, value) => {
  let arrayPath = [];

  // set initial L - M - R
  let left = 0;
  let right = searchList.length - 1;
  let middle = Math.floor(right / 2);

  // if first comparison finds value
  if (searchList[middle] == value) {
    arrayPath.push(searchList[middle]);
    return arrayPath;
  }

  while (searchList[middle] !== value) {
    // add to output array
    arrayPath.push(searchList[middle]);

    // not found
    if (right < left) {
      return 'Value Not Found';
    }
    // value is in left or right portion of array
    // update L - M - R
    if (searchList[middle] > value) {
      right = middle - 1;
      middle = left + Math.floor((right - left) / 2);
    } else {
      left = middle + 1;
      middle = left + Math.floor((right - left) / 2);
    }

    // if found update output array and exit
    if (searchList[middle] == value) {
      arrayPath.push(searchList[middle]);

      break;
    }
  }
  return arrayPath;
};
```
