---
id: a3f503de51cfab748ff001aa
title: ペアワイズ
challengeType: 1
forumTopicId: 301617
dashedName: pairwise
---

# --description--

配列 `arr` が与えられたとき、第2の引数 `arg` と等しくなる要素のペアを見つけ、そのインデックスの合計を返してください。

同じ数値の要素を持つ複数のペアを使うことができますが、インデックスは異なる必要があります。各ペアは可能な限り最も小さいインデックスを使うべきです。一度使われた要素は、別の要素とペアを作るために再利用できません。例えば、`pairwise([1, 1, 2], 3)` はインデックス0の1を使ってペア `[2, 1]` を作ります。これは 0+2 < 1+2 だからです。

例えば `pairwise([7, 9, 11, 13, 15], 20)` は `6` を返します。合計が20になるペアは `[7, 13]` と `[9, 11]` です。次に、それらのインデックスと値を配列で示します。

<div style='margin-left: 2em;'>

| インデックス | 0 | 1 | 2  | 3  | 4  |
| ----------- | - | - | -- | -- | -- |
| 値          | 7 | 9 | 11 | 13 | 15 |

</div>

以下で対応するインデックスを足し合わせます。

<div style='margin-left: 2em;'>

7 + 13 = 20 → インデックス 0 + 3 = 3  
9 + 11 = 20 → インデックス 1 + 2 = 3  
3 + 3 = 6 → `6` を返す

</div>

# --hints--

`pairwise([1, 4, 2, 3, 0, 5], 7)` は 11 を返すべきです。

```js
assert.deepEqual(pairwise([1, 4, 2, 3, 0, 5], 7), 11);
```

`pairwise([1, 3, 2, 4], 4)` は 1 を返すべきです。

```js
assert.deepEqual(pairwise([1, 3, 2, 4], 4), 1);
```

`pairwise([1, 1, 1], 2)` は 1 を返すべきです。

```js
assert.deepEqual(pairwise([1, 1, 1], 2), 1);
```

`pairwise([0, 0, 0, 0, 1, 1], 1)` は 10 を返すべきです。

```js
assert.deepEqual(pairwise([0, 0, 0, 0, 1, 1], 1), 10);
```

`pairwise([], 100)` は 0 を返すべきです。

```js
assert.deepEqual(pairwise([], 100), 0);
```

# --seed--

## --seed-contents--

```js
function pairwise(arr, arg) {
  return arg;
}

pairwise([1,4,2,3,0,5], 7);
```

# --solutions--

```js
function pairwise(arr, arg) {
  var sum = 0;
  arr.forEach(function(e, i, a) {
    if (e != null) {
      var diff = arg-e;
      a[i] = null;
      var dix = a.indexOf(diff);
      if (dix !== -1) {
        sum += dix;
        sum += i;
        a[dix] = null;
      }
    }
  });
  return sum;
}

pairwise([1,4,2,3,0,5], 7);
```
