---
id: 68a47dcdbce12193caf77012
title: 정렬된 인덱스 찾기 함수 구현하기
challengeType: 26
dashedName: lab-sorted-index-finder
---

# --description--

이 실습에서는 오름차순으로 정렬된 배열에 값을 삽입할 때 가장 낮은 인덱스를 반환하는 함수를 만드세요.

**목표**: 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `getIndexToIns` 함수는 배열과 숫자 두 개의 인수를 받아야 합니다.  
2. `sort` 메서드를 사용해 배열을 오름차순으로 정렬해야 합니다.
3. `getIndexToIns` 함수는 `findIndex` 메서드를 사용해 숫자가 삽입되어야 할 가장 낮은 인덱스를 반환해야 합니다.
4. `getIndexToIns` 함수는 항상 숫자를 반환해야 합니다.

**힌트:**  
`findIndex` 메서드는 JavaScript 내장 배열 메서드입니다. 콜백 함수를 받아 조건을 만족하는 첫 번째 요소의 인덱스를 반환합니다. `findIndex`와 `sort` 모두 고차 함수입니다.

**예시:**

- `getIndexToIns([1, 2, 3, 4], 1.5)`는 `1`를 반환해야 합니다. `1.5`가 `1` (인덱스 `0`)보다 크고 `2` (인덱스 `1`)보다 작기 때문입니다.
- `getIndexToIns([20, 3, 5], 19)`는 `2`를 반환해야 합니다. `[3, 5, 20]`로 정렬한 후 `19`가 `20` (인덱스 `2`)보다 작고 `5` (인덱스 `1`)보다 크기 때문입니다.


# --hints--

`getIndexToIns` 함수를 만들어야 합니다.

```js
assert.isFunction(getIndexToIns);
```

`getIndexToIns`는 항상 숫자를 반환해야 합니다.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

함수는 `sort` 메서드를 사용해야 합니다.

```js
assert.match(getIndexToIns.toString(), /\.sort\(/);
```


함수는 `findIndex` 메서드를 활용해야 합니다.

```js
assert.match(getIndexToIns.toString(), /\.findIndex\(/);
```

`getIndexToIns([10, 20, 30, 40, 50], 35)`는 `3`를 반환해야 합니다.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

`getIndexToIns([10, 20, 30, 40, 50], 30)`는 `2`를 반환해야 합니다.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

`getIndexToIns([40, 60], 50)`는 `1`를 반환해야 합니다.

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

`getIndexToIns([3, 10, 5], 3)`는 `0`를 반환해야 합니다.

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

`getIndexToIns([5, 3, 20, 3], 5)`는 `2`를 반환해야 합니다.

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

`getIndexToIns([2, 20, 10], 19)`는 `2`를 반환해야 합니다.

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

`getIndexToIns([3, 10, 5], 11)`는 `3`를 반환해야 합니다.

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 11), 3);
```

`getIndexToIns([], 5)`는 `0`를 반환해야 합니다.

```js
assert.strictEqual(getIndexToIns([], 5), 0);
```

# --seed--

## --seed-contents--

```js
``` 

# --solutions--

```js
function getIndexToIns(arr, num) {
  const sorted = [...arr].sort((a, b) => a - b);
  const index = sorted.findIndex(el => el >= num);

  return index === -1 ? sorted.length : index;
}
```
