---
id: 587d78b1367417b2b2512b0a
title: 為較高解析度顯示器使用 Retina 影像
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/cVZ4Rfp'
forumTopicId: 301142
dashedName: use-a-retina-image-for-higher-resolution-displays
---

# --description--

隨著連接網際網路的裝置數量增加，其尺寸和規格各異，且它們使用的顯示器在外部和內部可能不同。像素密度是某一裝置與其他裝置可能不同的一個面向，這種密度稱為 Pixel Per Inch(PPI) 或 Dots Per Inch(DPI)。最著名的此類顯示器是最新 Apple MacBook Pro 筆記型電腦以及近期 iMac 電腦上的「Retina Display」。由於「Retina」與「Non-Retina」顯示器之間的像素密度差異，某些未針對高解析度顯示器製作的影像，在高解析度顯示器上呈現時可能會看起來「像素化」。

讓你的影像能夠正確顯示在高解析度顯示器上（例如 MacBook Pro 的「retina display」）最簡單的方法，是將它們的 `width` 和 `height` 值定義為原始檔案的一半。以下是一個只使用原始高度和寬度一半的影像範例：

```html
<style>
  img {
    height: 250px;
    width: 250px;
  }
</style>
<img src="coolPic500x500" alt="A most excellent picture" />
```

# --instructions--

將 `img` 頁籤的 `width` 和 `height` 設定為原始值的一半。在此情況下，原始的 `height` 和原始的 `width` 均為 `200px`。

# --hints--

你的 `img` 頁籤應該有 100 像素的 `width`。

```js
assert.strictEqual(document.querySelector('img').width, 100);
```

你的 `img` 頁籤應該有 100 像素的 `height`。

```js
assert.strictEqual(document.querySelector('img').height, 100);
```

# --seed--

## --seed-contents--

```html
<style></style>

<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickers-CamperBot200x200.jpg"
  alt="freeCodeCamp sticker that says 'Because CamperBot Cares'"
/>
```

# --solutions--

```html
<style>
  img {
    height: 100px;
    width: 100px;
  }
</style>

<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickers-CamperBot200x200.jpg"
  alt="freeCodeCamp sticker that says 'Because CamperBot Cares'"
/>
```
