---
id: 587d78b1367417b2b2512b09
title: 使影像響應式
challengeType: 0
forumTopicId: 301140
dashedName: make-an-image-responsive
---

# --description--

使用 CSS 讓影像響應式其實非常簡單。你只需要為影像添加這些屬性：

```css
img {
  max-width: 100%;
  height: auto;
}
```

`max-width` 設為 `100%` 將確保圖片寬度不會超過其所在的容器，而 `height` 設為 `auto` 則會讓圖片保持原始的長寬比。

# --instructions--

為你的 `responsive-img` 類別添加樣式規則，使其具有響應式。它的寬度絕不應超過其容器（在此案例中為預覽視窗），且應保持其原始長寬比。當你添加完程式碼後，調整預覽視窗大小以查看你的圖片行為。

# --hints--

你的 `responsive-img` 類別應該設定 `max-width` 為 `100%`。

```js
assert.strictEqual(getComputedStyle(document.querySelector('.responsive-img')).maxWidth, '100%');
```

你的 `responsive-img` 類別應該將 `height` 設定為 `auto`。

```js
assert.match(code, /height:\s*?auto;/g);
```

# --seed--

## --seed-contents--

```html
<style>
.responsive-img {


}

img {
  width: 600px;
}
</style>

<img
  class="responsive-img"
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
```

# --solutions--

```html
<style>
  .responsive-img {
    max-width: 100%;
    height: auto;
  }

  img {
    width: 600px;
  }
</style>

<img
  class="responsive-img"
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
```
