---
id: 587d78b1367417b2b2512b0c
title: 使排版具響應式
challengeType: 0
forumTopicId: 301141
dashedName: make-typography-responsive
---

# --description--

你可以使用視窗單位來做響應式排版，而不是使用 `em` 或 `px` 來調整文字大小。視窗單位像百分比一樣是相對單位，但它們是基於不同的項目。視窗單位是相對於裝置的視窗尺寸（寬度或高度），而百分比是相對於父容器元素的大小。

四種不同的 viewport 單位是：

<ul><li><code>vw</code>（視窗寬度）：<code>10vw</code> 表示視窗寬度的 10%。</li><li><code>vh</code>（視窗高度）：<code>3vh</code> 表示視窗高度的 3%。</li><li><code>vmin</code>（視窗最小值）：<code>70vmin</code> 表示視窗較小尺寸（高度或寬度）的 70%。</li><li><code>vmax</code>（視窗最大值）：<code>100vmax</code> 表示視窗較大尺寸（高度或寬度）的 100%。</li></ul>

這是一個範例，將 `body` 頁籤設定為視窗寬度的 30%。

```css
body {
  width: 30vw;
}
```

# --instructions--

將 `h2` 標記的 `width` 設為視窗寬度的 80％，並將段落的 `width` 設為視窗較小尺寸的 75％。

# --hints--

你的 `h2` 頁籤應該有 80vw 的 `width`。

```js
assert.match(
  __helpers.removeCssComments(code),
  /h2\s*?{\s*?width:\s*?80vw;\s*?}/g
);
```

你的 `p` 頁籤應該有 75vmin 的 `width`。

```js
assert.match(
  __helpers.removeCssComments(code),
  /p\s*?{\s*?width:\s*?75vmin;\s*?}/g
);
```

# --seed--

## --seed-contents--

```html
<style>

</style>

<h2>Importantus Ipsum</h2>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```

# --solutions--

```html
<style>
  h2 {
    width: 80vw;
  }
  p {
    width: 75vmin;
  }
</style>

<h2>Importantus Ipsum</h2>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```
