---
id: 587d78b0367417b2b2512b08
title: 創建媒體查詢
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/cqwKrtm'
forumTopicId: 301139
dashedName: create-a-media-query
---

# --description--

媒體查詢是 CSS3 中介紹的一種新技術，根據不同的視窗大小改變內容的呈現方式。視窗是使用者在網頁上的可見區域，並且會根據用來存取網站的裝置而有所不同。

媒體查詢由媒體型別組成，如果該媒體型別符合文件顯示的裝置型別，則會套用樣式。你可以在你的媒體查詢中放入任意多的選擇器和樣式。

以下是一個媒體查詢的範例，當裝置的寬度小於或相等於 `100px` 時會傳回內容：

```css
@media (max-width: 100px) {
  /* CSS Rules */
}
```

且以下媒體查詢會在裝置高度大於或相等於 `350px` 時傳回內容：

```css
@media (min-height: 350px) {
  /* CSS Rules */
}
```

請記住，只有當媒體類型與所使用的裝置媒體類型相符時，媒體查詢內的 CSS 才會被套用。

# --instructions--

新增媒體查詢，使當裝置高度小於或相等於 `800px` 時，`p` 頁籤的 `font-size` 為 `10px`。

# --hints--

你應該為 `height` 小於或相等於 `800px` 的裝置宣告 `@media` 查詢。

```js
const media = new __helpers.CSSHelp(document).getCSSRules('media');
assert.isTrue(
  media.some(x => x.media?.mediaText?.includes('(max-height: 800px)'))
);
```

當裝置的 `height` 小於或相等於 `800px` 時，你的 `p` 元素應該有 `10px` 的 `font-size`。

```js
const rules = new __helpers.CSSHelp(document).getRuleListsWithinMedia(
  '(max-height: 800px)'
);
assert.strictEqual(
  rules?.find(x => x.selectorText === 'p')?.style?.fontSize,
  '10px'
);
```

當裝置的 `height` 超過 `800px` 時，你的 `p` 元素應該有初始的 `font-size` 為 `20px`。

```js
const ifPFirst = new __helpers.CSSHelp(document)
  .getCSSRules()
  ?.find(x => x?.selectorText === 'p' || x?.media);
assert.strictEqual(ifPFirst?.style?.fontSize, '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  p {
    font-size: 20px;
  }

  /* Only change code below this line */

  /* Only change code above this line */
</style>

<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```

# --solutions--

```html
<style>
  p {
    font-size: 20px;
  }

  @media (max-height: 800px) {
    p {
      font-size: 10px;
    }
  }
</style>

<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```
