---
id: 69a5f35669099ed52f8563b1
title: بناء برنامج ذكي لإعادة تخزين المخزن
challengeType: 26
dashedName: lab-smart-pantry-restocker
---

# --description--

في هذا المختبر، ستبني برنامجًا صغيرًا لإدارة المخزن باستخدام مفاهيم JavaScript الأساسية مثل المصفوفات، والكائنات، والحلقات، والشرطيات.

ستحاكي استقبال شحنة من مواد المخزن، وتقرر ما يجب فعله مع كل مادة، وتنظم النتائج للتخزين.

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

تحتوي مصفوفة `rawData` على سلاسل نصية مفصولة بعلامة أنبوب بالصيغة `sku|name|qty|expires|zone`، حيث أن `zone` اختياري.

**قصص المستخدم:**

1. يجب أن تنفذ دالة `parseShipment(rawData)` التي تأخذ مصفوفة من السلاسل النصية وتُرجع مصفوفة من الكائنات التي تحتوي على الخصائص `{ sku, name, qty, expires, zone }`.

- يجب تجاهل قيم `sku` المكررة في الشحنة.
   - عندما لا يُذكر جزء `zone`، يجب أن يكون الافتراضي `"general"`.
   - يجب تحويل قيمة `qty` إلى رقم.

2. يجب أن تنفذ دالة `planRestock(pantry, shipment)` التي تقارن المخزن الحالي مع الشحنة الواردة وتُرجع مصفوفة من الإجراءات على شكل `{ type, item }`، حيث يكون `type` واحدًا من `"restock"`، `"discard"`، أو `"donate"`، و`item` هو كائن الشحنة المحلل.

معلمة `pantry` هي مصفوفة من الكائنات بنفس شكل كائن الشحنة المحلل (`{ sku, name, qty, expires, zone }`).

- إذا كان عنصر الشحنة يحتوي على قيمة `qty` تساوي `0` أو أقل، يجب أن يكون نوع الإجراء `type` هو `"discard"` بغض النظر عما إذا كان العنصر موجودًا في المخزن أم لا.
   - وإلا، إذا كان `sku` الخاص بالعنصر موجودًا بالفعل في المخزن، يجب أن يكون نوع الإجراء `type` هو `"restock"`.
   - وإلا (إذا لم يكن `sku` الخاص بالعنصر موجودًا في المخزن)، يجب أن يكون نوع الإجراء `type` هو `"donate"`.

3. يجب أن تنفذ دالة `groupByZone(actions)` التي تجمع الإجراءات حسب مناطق التخزين بناءً على خاصية `zone` لكل عنصر. يجب أن تُرجع الدالة كائنًا حيث يكون كل مفتاح اسم منطقة والقيمة مصفوفة من الإجراءات التابعة لتلك المنطقة. على سبيل المثال، إذا احتوت الإجراءات على عناصر بمناطق `"fridge"` و`"pantry"`، يجب أن تكون النتيجة `{ fridge: [...], pantry: [...] }`.

4. يجب أن تنفذ دالة `clonePantry(pantry)` التي تُرجع نسخة عميقة من المخزن حتى لا تؤثر تغييرات التخطيط على القائمة الأصلية. النسخة العميقة تعني إنشاء مصفوفة جديدة بكائنات جديدة، بحيث لا يغير تعديل النسخة المخزن الأصلي.

5. يجب أن تستخدم جميع الدوال معًا لمعالجة شحنة وتسجيل كائن النتيجة النهائي المجمع في وحدة التحكم.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن تعرف دالة باسم `parseShipment` تقبل مصفوفة واحدة من السلاسل النصية تسمى `rawData` كمعلمة.

``` js
assert.isFunction(parseShipment);
assert.lengthOf(parseShipment, 1);
```

يجب أن تحول دالة `parseShipment` سلاسل الشحنة في المصفوفة إلى كائنات تحتوي على الخصائص: `sku`، `name`، `qty`، `expires`، و`zone`.

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);

assert.isArray(result);
assert.isObject(result[0]);
assert.deepEqual(result[0], {
  sku: "A10",
  name: "Tomatoes",
  qty: 5,
  expires: "2027-01-01",
  zone: "fridge"
});
```

يجب تجاهل قيم SKU المكررة في الشحنة.

``` js
const rawData = [
  "C32|Eggs|3|2027-01-01",
  "C32|Eggs|3|2027-01-01"
];

const result = parseShipment(rawData);

assert.lengthOf(result, 1);
assert.strictEqual(result[0].sku, "C32");
```

عندما لا تكون `zone` موجودة في سلسلة الشحنة، يجب تعيينها إلى `general`.

``` js
const rawData = ["B21|Bananas|10|2027-01-01"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].zone, "general");
```

يجب تحويل قيمة `qty` إلى رقم.

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].qty, 5);
assert.isNumber(result[0].qty);
```

يجب أن تعرف دالة باسم `planRestock` تقبل معلمتين: `pantry` و`shipment`.

``` js
assert.isFunction(planRestock);
assert.lengthOf(planRestock, 2);
```

يجب أن تُرجع دالة `planRestock` مصفوفة من الإجراءات التي تحتوي على خصائص `type` و`item`.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 6, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.isArray(result);
assert.isObject(result[0]);
assert.property(result[0], "type");
assert.property(result[0], "item");
```

العناصر التي تكون كميتها `0` أو أقل في الشحنة يجب أن تُنشئ إجراءً بنوع `discard`، بغض النظر عن وجود العنصر في المخزن.

``` js
const pantry = [{ sku: "D43", name: "Pineapples", qty: 2, expires: "2027-01-01", zone: "general" }];
const shipment = [{ sku: "D43", name: "Pineapples", qty: 0, expires: "2027-01-01", zone: "general" },
                  { sku: "E54", name: "Peppers", qty: -1, expires: "2027-01-01", zone: "fridge" }
                ];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "discard");
assert.strictEqual(result[1].type, "discard");
```

عندما يكون عنصر الشحنة موجودًا بالفعل في المخزن (بنفس `sku`)، يجب أن يكون الإجراء `restock`.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "restock");
```

إذا لم يكن عنصر الشحنة موجودًا في المخزن، يجب أن يكون الإجراء `"donate"`.

``` js
const pantry = [];
const shipment = [{ sku: "C32", name: "Eggs", qty: 3, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "donate");
```

يجب أن تعرف دالة باسم `groupByZone` تقبل معلمة واحدة تسمى `actions`.

``` js
assert.isFunction(groupByZone);
assert.lengthOf(groupByZone, 1);
```

يجب أن تُرجع دالة `groupByZone` الإجراءات مجمعة حسب خاصية `zone` لكل كائن.

``` js

const actions = [
  { type: "restock", item: { sku: "A1", zone: "fridge" } },
  { type: "restock", item: { sku: "B1", zone: "pantry" } }, 
];

const result = groupByZone(actions);

assert.isObject(result);
assert.isArray(result.fridge);
assert.isArray(result.pantry);
```

يجب أن تجمع دالة `groupByZone` الإجراءات بشكل صحيح مع المحتوى والعدد المناسبين.

``` js
const actions = [
  { type: "restock", item: { sku: "A1", zone: "fridge" } },
  { type: "donate",  item: { sku: "B1", zone: "fridge" } },
  { type: "discard", item: { sku: "C1", zone: "pantry" } },
];

const result = groupByZone(actions);

assert.lengthOf(result.fridge, 2);
assert.lengthOf(result.pantry, 1);
assert.strictEqual(result.fridge[0].item.sku, "A1");
assert.strictEqual(result.pantry[0].type, "discard");
```

يجب أن تعرف دالة باسم `clonePantry` تقبل معلمة واحدة تسمى `pantry`.

``` js
assert.isFunction(clonePantry);
assert.lengthOf(clonePantry, 1);
```

يجب أن تُرجع دالة `clonePantry` مصفوفة جديدة بدلًا من مصفوفة المخزن الأصلية.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.isArray(pantryClone);
assert.notStrictEqual(pantryClone, pantry);
```

يجب أن تكون الكائنات داخل المخزن المستنسخ أيضًا كائنات جديدة (نسخة عميقة).

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.notStrictEqual(pantryClone[0], pantry[0]);
assert.deepEqual(pantryClone[0], pantry[0]);

pantryClone[0].qty = 100;
assert.strictEqual(pantry[0].qty, 4);
```

يجب أن تعمل الدوال معًا لمعالجة شحنة وتجميع الإجراءات الناتجة.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];

const shipment = parseShipment(rawData);
const pantryClone = clonePantry(pantry);
const actions = planRestock(pantryClone, shipment);
const grouped = groupByZone(actions);

assert.isObject(grouped);
assert.strictEqual(grouped.fridge[0].item.qty, 5);
```

يجب أن تسجل الإجراءات الناتجة مجمعة حسب المناطق.

``` js
const groupedLog = getLogs().find(log => {
  if (!log || typeof log !== "object" || Array.isArray(log)) return false;

  return Object.values(log).some(zoneActions => {
    if (!Array.isArray(zoneActions)) return false;

    return zoneActions.some(action =>
      action &&
      ["restock", "discard", "donate"].includes(action.type) &&
      action.item &&
      typeof action.item.sku === "string"
    );
  });
});

assert.isOk(groupedLog, "Log the grouped actions object to the console.");
```

# --seed--

## --seed-contents--

```js
const pantry = [
  { sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" },
  { sku: "D43", name: "Pineapples", qty: 2, expires: "2020-01-01", zone: "general" }
];

const rawData = [
  "A10|Tomatoes|5|2027-01-01",
  "B21|Bananas|10|2027-01-01",
  "C32|Eggs|3|2027-01-01|fridge",
  "C32|Eggs|3|2027-01-01",
  "D43|Pineapples|0|2027-01-01",
  "E54|Peppers|-1|2027-01-01|fridge"
];
```

# --solutions--

```js

function parseShipment(rawData) {

  // Early return if rawData is not an Array
  if (!Array.isArray(rawData)) return [];

  const result = [];
  const seenSKUs = [];

  for (let line of rawData) {

    const parts = line.split("|");

    const sku = parts[0].trim();

    if (seenSKUs.includes(sku)) continue; // Avoid duplicates

    const name = parts[1].trim();
    const qty = Number(parts[2].trim()); // Parse qty into a number
    const expires = parts[3].trim();
    const zone = ( parts[4] ) ? parts[4].trim() : "general";

    result.push({
      sku: sku,
      name: name,
      qty: qty,
      expires: expires,
      zone: zone
    });

    seenSKUs.push(sku);

  }

  return result;
}

function planRestock(pantry, shipment) {
  const actions = [];

  for (let item of shipment) {

    var exists = false;

    for (let stock of pantry)
    {
      if (stock.sku === item.sku)
      {
        exists = true;
      }
    }

    if (item.qty > 0) {
      if (exists) {
        actions.push({ type: "restock", item: item });
      } else{
        actions.push({ type: "donate", item: item });
      }
    } else {
      actions.push({ type: "discard", item: item });
    }
  }

  return actions;
}

function groupByZone(actions) {

  const byZone = {};

  for (const action of actions) {

    const zone = action.item.zone;

    if (!byZone[zone]) {
      byZone[zone] = [];
    }

    byZone[zone].push(action);
  }

  return byZone;

}

function clonePantry(pantry) {
  const pantryClone = [];

  for (let item of pantry) {
    pantryClone.push({
      sku: item.sku,
      name: item.name,
      qty: item.qty,
      expires: item.expires,
      zone: item.zone
    });
  }

  return pantryClone;
}

const pantry = [
  { sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" },
  { sku: "D43", name: "Pineapples", qty: 2, expires: "2020-01-01", zone: "general" }
];

const rawData = [
  "A10|Tomatoes|5|2027-01-01", // Restock existing item
  "B21|Bananas|10|2027-01-01", // Donate new item without zone
  "C32|Eggs|3|2027-01-01|fridge", // Donate to a defined zone
  "C32|Eggs|3|2027-01-01", // Duplicated SKU in shipment
  "D43|Pineapples|0|2027-01-01", // Discard with quantity of 0
  "E54|Peppers|-1|2027-01-01|fridge" // Discard even if it's not in pantry
];

const shipment = parseShipment(rawData);
const pantryCopy = clonePantry(pantry);
const actions = planRestock(pantryCopy, shipment);
const grouped = groupByZone(actions);

// Show items expanded in zones
console.log(grouped);

```
