---
id: 68d275dd800f404d22a07564
title: "implementare l'algoritmo di ricerca in profondità"
challengeType: 27
dashedName: lab-depth-first-search
---

# --description--

In questo laboratorio implementerai un algoritmo di attraversamento di un grafo chiamato <dfn>ricerca in profondità</dfn>.

Mentre la <dfn>ricerca in ampiezza</dfn> esplora i nodi a distanze incrementali dal nodo sorgente, la <dfn>ricerca in profondità</dfn> scende lungo un percorso di archi finché può.

Una volta raggiunta una fine del percorso, la ricerca torna indietro all'ultimo nodo con un percorso di archi non visitato e continua la ricerca.

A differenza della ricerca in ampiezza, ogni volta che un nodo viene visitato, non visita tutti i suoi vicini. Invece, visita prima uno dei suoi vicini e continua lungo quel percorso finché non ci sono più nodi da visitare su quel percorso.

Per implementare questo algoritmo, vorrai usare uno stack (un array in cui l'ultimo elemento aggiunto è il primo a essere rimosso, seguendo il principio <dfn>Last-In-First-Out</dfn>). Uno stack è utile negli algoritmi di ricerca in profondità perché, mentre aggiungi i vicini allo stack, vuoi visitare prima i vicini aggiunti più di recente e rimuoverli dallo stack.

Un output semplice di questo algoritmo è una lista di nodi raggiungibili da un nodo dato. Perciò, vorrai anche tenere traccia dei nodi che visiti.

**Obiettivo**: Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione chiamata `dfs`.
1. La funzione `dfs` dovrebbe prendere due argomenti:
   - Una matrice di adiacenza non orientata.
   - Un'etichetta di nodo, che è il valore numerico del nodo tra `0` e `n - 1`, dove `n` è il numero totale di nodi nel grafo.
1. La funzione `dfs` dovrebbe implementare l'algoritmo di ricerca in profondità per restituire una lista di tutti i nodi raggiungibili dal nodo passato come argomento.

# --hints--

Dovresti avere una funzione chiamata `dfs` che prende due argomenti.

```js
({ test: () => runPython(`
  import inspect
  assert inspect.isfunction(dfs)
  sig = inspect.signature(dfs)
  assert len(sig.parameters) == 2
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)` dovrebbe restituire una lista con `1`, `2`, `3` e `0`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)
  assert isinstance(result, list)
  assert len(result) == 4
  assert set(result) == {1, 2, 3, 0}
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)` dovrebbe restituire una lista con `1`, `2`, `3` e `0`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)
  assert isinstance(result, list)
  assert len(result) == 4
  assert set(result) == {3, 2, 1, 0}
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)` dovrebbe restituire `[3]`.

```js
({ test: () => runPython(`
  assert dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3) == [3]
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)` dovrebbe restituire una lista con `3` e `2`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)
  assert result == [3, 2] or result == [2, 3]
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)` dovrebbe restituire una lista con `0` e `1`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)
  assert result == [0, 1] or result == [1, 0]
`) })
```

La funzione `dfs` dovrebbe restituire i risultati corretti.

```js
({ test: () => runPython(`
  import random

  def solve(g, s):
      v, q = set(), [s]
      while q:
          u = q.pop()
          if u not in v:
              v.add(u)
              q.extend([i for i, x in enumerate(g[u]) if x == 1 and i not in v])
      return v

  random.seed(0)

  for _ in range(10):
      n = 6
      mat = [[0]*n for _ in range(n)]
      for i in range(n):
          for j in range(i+1, n):
              if random.random() > 0.5:
                  mat[i][j] = mat[j][i] = 1

      start = random.randint(0, n-1)
      assert set(dfs(mat, start)) == solve(mat, start), "Random graph check failed"
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def dfs(graph, root):
    stack = []
    temp_v = None
    visited = []
    temp_v_neighbors = []
    stack.append(root)
    while stack:
        temp_v = stack.pop()
        if temp_v not in visited:
            visited.append(temp_v)
            temp_v_neighbors = graph[temp_v]
            for n, is_neighbor in enumerate(temp_v_neighbors):
                if is_neighbor == 1:
                    stack.append(n)
    return visited
```
