---
id: 613297a923965e0703b64796
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

قد تكون مألوفًا مع عنصر `meta` بالفعل؛ يُستخدم لتحديد معلومات عن الصفحة، مثل العنوان، الوصف، الكلمات المفتاحية، والمؤلف.

أضف إلى صفحتك عنصر `meta` مع قيمة `charset` مناسبة.

تحدد خاصية `charset` ترميز الأحرف للصفحة، وفي الوقت الحالي، `UTF-8` هو الترميز الوحيد الذي تدعمه معظم المتصفحات.

# --hints--

يجب إنشاء عنصر `meta` داخل عنصر `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

يجب أن تعطي وسم `meta` خاصية `charset` بقيمة `UTF-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    --fcc-editable-region--
    
    --fcc-editable-region--
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
