---
id: 6133d11ef548f51f876149e3
title: خطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

التنقل جزء أساسي من إمكانية الوصول، ويعتمد قارئ الشاشة على توفير هيكل الصفحة. يتم ذلك باستخدام عناصر HTML الدلالية.

أضف عنصر `header` وعنصر `main` إلى صفحتك.

يُستخدم عنصر `header` لتقديم الصفحة، بالإضافة إلى توفير قائمة تنقل.

يحتوي عنصر `main` على المحتوى الأساسي لصفحتك.

# --hints--

يجب إضافة عنصر `header` إلى عنصر `body`.

```js
assert.exists(document.querySelector('body > header'));
```

يجب إضافة عنصر `main` إلى عنصر `body`.

```js
assert.exists(document.querySelector('body > main'));
```

يجب أن يأتي عنصر `header` قبل عنصر `main`.

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
