---
id: 6133acc353338c0bba9cb553
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

أخيرًا في العنصر `head`، عنصر `title` مفيد لقارئ الشاشة لفهم محتوى الصفحة. بالإضافة إلى ذلك، هو جزء مهم من _SEO_.

امنح صفحتك `title` يكون وصفيًا ومختصرًا.

# --hints--

يجب إضافة عنصر `title` إلى العنصر `head`.

```js
assert.exists(document.querySelector('head > title'));
```

يجب ألا يكون `title` أطول من 60 حرفًا.

```js
assert.isAtMost(document.querySelector('head > title')?.textContent?.length, 60);
```

حاول أن تجعل عنصر `title` أكثر وصفًا. _تلميح: على الأقل 15 حرفًا_

```js
assert.isAtLeast(document.querySelector('head > title')?.textContent?.length, 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    --fcc-editable-region--
    
    --fcc-editable-region--
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
