---
id: 613e2546d0594208229ada50
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

ضمن الـ `header`، قدّم سياقًا عن الصفحة عبر تسكين عنصر واحد من كل من `img` و `h1` و `nav`.

يجب أن يشير الـ `img` إلى `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`، وأن يحمل خاصية `id` بقيمة `logo`، وأن يحتوي على نص بديل `alt` بقيمة `freeCodeCamp`.

يجب أن يحتوي الـ `h1` على النص `HTML/CSS Quiz`.

# --hints--

أضف عنصر `img` إلى عنصر الـ `header`.

```js
assert.exists(document.querySelector('header > img'));
```

أضف عنصر `h1` إلى عنصر الـ `header`.

```js
assert.exists(document.querySelector('header > h1'));
```

أضف عنصر `nav` إلى عنصر الـ `header`.

```js
assert.exists(document.querySelector('header > nav'));
```

ضع عناصر `img` و `h1` و `nav` بهذا الترتيب.

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

امنح عنصر الـ `img` خاصية `src` بقيمة `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

امنح عنصر الـ `img` خاصية `id` بقيمة `logo`.

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

امنح عنصر الـ `img` خاصية `alt` بقيمة `freeCodeCamp`.

```js
assert.equal(document.querySelector('img')?.alt, 'freeCodeCamp');
```

امنح عنصر الـ `h1` النص `HTML/CSS Quiz`.

```js
assert.include(document.querySelector('h1')?.innerText.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
