---
id: 6140827cff96e906bd38fc2b
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

كما هو موضح في <a href="https://design-style-guide.freecodecamp.org/" target="_blank" rel="noopener noreferrer nofollow">دليل أسلوب freeCodeCamp</a>، يجب أن يحتفظ الشعار بنسبة أبعاد `35 / 4`، وأن يكون له حشو حول النص.

أولًا، غيّر `background-color` إلى `#0a0a23` حتى ترى الشعار. ثم استخدم خاصية `aspect-ratio` لتعيين نسبة الأبعاد المطلوبة إلى `35 / 4`. أخيرًا، أضف حشو `padding` بقيمة `0.4rem` حول الشعار.

# --hints--

يجب أن تعطي `#logo` خاصية `background-color` بقيمة `#0a0a23`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.backgroundColor, 'rgb(10, 10, 35)');
```

يجب أن تستخدم خاصية `aspect-ratio`.

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio);
```

يجب ألا تستخدم خاصية `height`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.height);
```

يجب ألا تغيّر خاصية `width`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

يجب أن تعطي عنصر `img` خاصية `aspect-ratio` بقيمة `35 / 4`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio, '35 / 4');
```

يجب أن تعطي عنصر `img` حشو `padding` بقيمة `0.4rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.padding, '0.4rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

#logo {
  width: max(10rem, 18vw);
  --fcc-editable-region--
  
  --fcc-editable-region--
}

```
