---
id: 614202874ca576084fca625f
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

كل دور `region` يتطلب مسمى، يساعد مستخدمي قارئ الشاشة على فهم هدف المنطقة. إحدى الطرق لإضافة مسمى هي إضافة عنصر عنوان داخل المنطقة ثم الإشارة إليه باستخدام الخاصية `aria-labelledby`.

أضف الخاصيات `aria-labelledby` التالية إلى عناصر `section`:

- `student-info`
- `html-questions`
- `css-questions`

ثم، داخل كل عنصر `section`، قم بتسكين عنصر `h2` واحد مع خاصية `id` تطابق الخاصية `aria-labelledby` المقابلة. أعطِ كل `h2` محتوى نصيًا مناسبًا.

# --hints--

يجب أن تعطي عنصر `section` الأول خاصية `aria-labelledby` بقيمة `student-info`.

```js
assert.equal(document.querySelectorAll('section')?.[0]?.getAttribute('aria-labelledby'), 'student-info');
```

يجب أن تعطي عنصر `section` الثاني خاصية `aria-labelledby` بقيمة `html-questions`.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.getAttribute('aria-labelledby'), 'html-questions');
```

يجب أن تعطي عنصر `section` الثالث خاصية `aria-labelledby` بقيمة `css-questions`.

```js
assert.equal(document.querySelectorAll('section')?.[2]?.getAttribute('aria-labelledby'), 'css-questions');
```

يجب أن تسكن عنصر `h2` واحد داخل عنصر `section` الأول.

```js
assert.equal(document.querySelectorAll('section')?.[0]?.querySelectorAll('h2')?.length, 1);
```

يجب أن تسكن عنصر `h2` واحد داخل عنصر `section` الثاني.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.querySelectorAll('h2')?.length, 1);
```

يجب أن تسكن عنصر `h2` واحد داخل عنصر `section` الثالث.

```js
assert.equal(document.querySelectorAll('section')?.[2]?.querySelectorAll('h2')?.length, 1);
```

يجب أن تعطي عنصر `h2` الأول خاصية `id` بقيمة `student-info`.

```js
assert.equal(document.querySelectorAll('h2')?.[0]?.id, 'student-info');
```

يجب أن تعطي عنصر `h2` الثاني خاصية `id` بقيمة `html-questions`.

```js
assert.equal(document.querySelectorAll('h2')?.[1]?.id, 'html-questions');
```

يجب أن تعطي عنصر `h2` الثالث خاصية `id` بقيمة `css-questions`.

```js
assert.equal(document.querySelectorAll('h2')?.[2]?.id, 'css-questions');
```

يجب أن تعطي عنصر `h2` الأول محتوى نصيًا مناسبًا. _تلميح: كنت سأختار `Student Info`_

```js
assert.isNotEmpty(document.querySelectorAll('h2')[0]?.textContent.trim());
```

يجب أن تعطي عنصر `h2` الثاني محتوى نصيًا مناسبًا. _تلميح: كنت سأختار `HTML`_

```js
assert.isNotEmpty(document.querySelectorAll('h2')[1]?.textContent.trim());
```

يجب أن تعطي عنصر `h2` الثالث محتوى نصيًا مناسبًا. _تلميح: كنت سأختار `CSS`_

```js
assert.isNotEmpty(document.querySelectorAll('h2')[2]?.textContent.trim());
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        --fcc-editable-region--
        <section role="region">
          
        </section>
        <section role="region">
          
        </section>
        <section role="region">
          
        </section>
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
