---
id: 614ccc21ea91ef1736b9b578
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

لقد زودت في هذه الورشة بقالب جاهز أساسي.

ابدأ رحلة إمكانية الوصول هذه بإضافة خاصية `lang` إلى عنصر `html`. هذا يساعد قارئ الشاشة على تحديد لغة الصفحة.

تذكر أيضًا ربط ورقة الأنماط بالصفحة.

# --hints--

يجب أن تعطي عنصر `html` خاصية `lang`. _تلميح: يمكنك استخدام القيمة `en` للإنجليزية._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

يجب أن يحتوي عنصر `head` على عنصر `link`.

```js
assert.exists(document.querySelector("head>link"));
```

يجب أن يحتوي عنصر `link` على `rel="stylesheet"` و `href="styles.css"`.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>
    
--fcc-editable-region--
  </head>
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
