---
id: 68507cadbf36aa089a84ce2d
title: "implementare l'algoritmo di Luhn"
challengeType: 27
dashedName: implement-the-luhn-algorithm
---

# --description--

L'algoritmo di Luhn, noto anche come algoritmo "modulo 10" o "mod 10", è una formula di checksum semplice usata per validare una varietà di numeri identificativi, come i numeri delle carte di credito. Questi sono i passaggi per validare un numero usando l'algoritmo di Luhn:

- Partendo da destra, ed escludendo la cifra più a destra (la cifra di controllo), raddoppia il valore di ogni altra cifra.
- Se il risultato del raddoppio di una cifra è maggiore di `9`, somma le cifre per ottenere una cifra singola. In alternativa, puoi sottrarre `9` dal risultato.
- Prendi la somma di tutte le cifre inclusa la cifra di controllo.
- Se la somma di tutte le cifre è un multiplo di `10`, allora il numero è valido; altrimenti non è valido.

Supponiamo di avere il numero `453914881`. I passaggi per validarlo usando l'algoritmo di Luhn sarebbero:

 ```md 
 Numero conto       4   5   3   9   1   4   8   8   1 
 Raddoppia ogni altro  4  10   3  18   1   8   8  16   1 
 Somma cifre a 2 caratteri   4   1   3   9   1   8   8   7   1 
 ``` 

- Poi somma tutti i numeri, `4 + 1 + 3 + 9 + 1 + 8 + 8 + 7 + 1 = 42`.
- Poiché `42` non è un multiplo di `10`, il numero non è valido.

In questo laboratorio, costruirai un validatore di carte di credito usando l'algoritmo di Luhn.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti definire una funzione chiamata `verify_card_number` che prende una stringa di cifre (che rappresenta un numero di carta) e verifica se è valida secondo l'algoritmo di Luhn.

2. All'interno della funzione `verify_card_number`:

   - Dovresti gestire eventuali trattini o spazi presenti nel numero di carta passato.
   - Restituisci `VALID!` se il numero di carta è valido; altrimenti, restituisci `INVALID!`.

Quando completi il progetto, dovresti vedere i seguenti messaggi a seconda dell'input:

| Numero carta         | Messaggio  |
| ------------------- | -------- |
| 453914889           | VALID!   |
| 4111-1111-1111-1111 | VALID!   |
| 1234 5678 9012 3456 | INVALID! |


# --hints--

Dovresti avere una funzione chiamata `verify_card_number`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('verify_card_number')
    `)
})
```

`verify_card_number('453914889')` dovrebbe restituire `VALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914889') == 'VALID!'
    `)
})
```

`verify_card_number('4111-1111-1111-1111')` dovrebbe restituire `VALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('4111-1111-1111-1111') == 'VALID!'
    `)
})
```

`verify_card_number('453914881')` dovrebbe restituire `INVALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914881') == 'INVALID!'
    `)
})
```

`verify_card_number('1234 5678 9012 3456')` dovrebbe restituire `INVALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('1234 5678 9012 3456') == 'INVALID!'
    `)
})
```

`verify_card_number` dovrebbe restituire `VALID!` quando chiamata con un numero di carta di credito valido.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6467') == 'VALID!'
    assert verify_card_number('6011 1111 1111 1117') == 'VALID!'
    assert verify_card_number('3782-822463-10005') == 'VALID!'
    assert verify_card_number('5555 5555 5555 4444') == 'VALID!'
    assert verify_card_number('3056 9309 0259 04') == 'VALID!'
    assert verify_card_number('6011000990139424173') == 'VALID!'
    `)
})
```

`verify_card_number` dovrebbe restituire `INVALID!` quando chiamata con un numero di carta di credito non valido.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6466') == 'INVALID!'
    assert verify_card_number('6011 1111 1111 1116') == 'INVALID!'
    assert verify_card_number('3782 822463 10004') == 'INVALID!'
    assert verify_card_number('5555-5555-5555-4440') == 'INVALID!'
    assert verify_card_number('3056 9309 0259 05') == 'INVALID!'
    assert verify_card_number('6011000990139424175') == 'INVALID!'
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def verify_card_number(card_number):
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    sum_of_odd_digits = 0
    card_number_reversed = translated_card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            number = (number // 10) + (number % 10)
        sum_of_even_digits += number

    total = sum_of_odd_digits + sum_of_even_digits
    if total % 10 == 0:
        return 'VALID!'
    else:
        return 'INVALID!'

# Example usage
verify_card_number('4111-1111-4555-1142')
```
