---
id: 6852f11d412c29665985a915
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

既然使用者可以發送電子郵件，他們需要一種方式來儲存收到的電子郵件。每個使用者都應該有一個收件匣來收集他們的電子郵件。

為你的 `User` 類別的 `__init__` 方法添加一個 `inbox` 屬性，並將其初始化為空列表 `[]`。這將用來儲存使用者收到的所有電子郵件。

# --hints--

你應該將 `self.inbox = []` 添加到 `User` 類別的 `__init__` 方法中。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
```
