---
id: 6852f01116912861c7d1b0e6
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

在 `__init__` 方法中，你需要將參數值指定到物件的屬性，以便每封電子郵件都能儲存其資訊。

在你的 `__init__` 方法中，移除 `pass` 關鍵字，並將參數指定到 `self.sender`、`self.receiver`、`self.subject` 和 `self.body`。

# --hints--

你應該將 `sender` 參數指定給 `self.sender`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.sender").is_equivalent("self.sender = sender")`))
    }

})
```

你應該將 `receiver` 參數指定給 `self.receiver`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.receiver").is_equivalent("self.receiver = receiver")`))
    }

})
```

你應該將 `subject` 參數指定給 `self.subject`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.subject").is_equivalent("self.subject = subject")`))
    }

})
```

你應該將 `body` 參數指定給 `self.body`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.body").is_equivalent("self.body = body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        pass
--fcc-editable-region--
```
