---
id: 6852ef95ccc27660275692cd
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

你的 `Email` 類別需要在創建時儲存每封電子郵件的資訊。

在你的 `Email` 類別中，移除現有的 `pass` 關鍵字並創建接受 `self`、`sender`、`receiver`、`subject` 和 `body` 作為參數的 `__init__` 方法。

# --hints--

你應該在 `Email` 類別內定義一個 `__init__` 方法。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__init__")`))
    }

})
```

你的 `__init__` 方法應該有參數：`self`、`sender`、`receiver`、`subject` 和 `body`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").has_args("self, sender, receiver, subject, body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    pass
--fcc-editable-region--
```
