---
id: 6852f06fdde1b4634a263f03
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

現在讓我們透過建立一個 `Email` 類別的電子郵件物件並檢查它是否正確儲存資訊來測試它。你將列印幾個屬性以驗證一切是否正常運作。

創建一個名為 `email_obj` 的電子郵件物件，將 `alice@example.com` 設為寄件者，`bob@example.com` 設為收件者，`Hello` 設為主旨，`Hi Bob!` 設為內容。然後分別列印 `sender` 和 `subject` 屬性作為獨立的列印述語，以驗證它們是否正確儲存。

# --hints--

你應該使用指定的參數創建一個名為 `email_obj` 的電子郵件物件。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_variable("email_obj").is_equivalent("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`))
    }

})
```

你應該列印 `email_obj.sender`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.sender)")`))
    }

})
```

你應該列印 `email_obj.subject`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.subject)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--

--fcc-editable-region--
```
