---
id: 6852ea0c3cadb54e897acdac
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個工作坊中，你將建置一個*Email Simulator*，模擬不同使用者之間發送、接收和管理電子郵件。你將學習類別、物件，以及如何以物件導向的方式組織程式碼。

首先使用 `class` 關鍵字建立一個名為 `Email` 的類別。

# --hints--

你應該創建一個名為 `Email` 的類別。

```js
({
    test: () => 
    {
        assert(runPython(`(_Node(_code).has_class("Email"))`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
