---
id: 6852f0c3fcedb964ab959921
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

現在讓我們測試 `read` 屬性是否已正確添加到你的 `Email` 類別。由於你已經有來自先前步驟的 email 物件，列印 `read` 屬性以查看它現在的預設值是否為 `False`。

# --hints--

你應該列印 `email_obj.read` 來測試讀取屬性。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)

--fcc-editable-region--
```
