---
id: 6853f4ded9f84896ff182dd8
title: 步驟 31
challengeType: 20
dashedName: step-31
---

# --description--

在空收件匣檢查後，列印訊息 `\nYour Emails:`。之後，使用帶有列舉的 `for` 迴圈迭代收件匣中的所有電子郵件。這是帶有起始數字的列舉語法：`enumerate(iterable, start=start_number)`。

使用列舉從 `1` 開始為電子郵件編號。使用 `i`（索引）和 `email`（電子郵件物件）作為迭代變數。

在迴圈內，列印一個 f-string，內容為迭代索引後接一個 `.`，然後是一個空格及以此格式表述的 email 物件字串：

```md
i. string_representation
```

其中 `i` 是索引，`string_representation` 是 email 物件的表述。

# --hints--

你應該在 `if` 述語外列印標題 `\nYour Emails:`。

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_body()[1].is_equivalent("print('\\\\nYour Emails:')")`))
    }
})
```

你應該有一個 `for` 迴圈，迭代 `enumerate(self.emails, start=1)`。

```js
({
    test: () => 
    {
        runPython(`
_iter = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_iter()
assert _iter.is_equivalent("enumerate(self.emails, start=1)") or _iter.is_equivalent("enumerate(self.emails, 1)")
`)
    }
})
```

你的 `for` 迴圈應該使用 `i` 和 `email` 來迭代 `enumerate(self.emails, start=1)`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_vars().is_equivalent("i, email")`))
    }
})
```

在迴圈中，你應該列印每封電子郵件摘要，前置式為編號，後面接著來自 `__str__` 方法的格式化字串。別忘了在編號後面添加 `.`。

```js
({
    test: () => 
    {
        assert(runPython(`
        cond = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_body()
        cond.is_equivalent("print(f'{i}. {str(email)}')") or cond.is_equivalent("print(f'{i}. {email}')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        
--fcc-editable-region--
```
