---
id: 6853f2da9b8ed891e9fdc5fc
title: 第 23 步
challengeType: 20
dashedName: step-23
---

# --description--

現在讓我們顯示寄件者和收件者的資訊。若要顯示誰寄送和接收電子郵件，請為 `display_full_email` 方法添加兩個列印述語，格式如下：

- `From: sender` 其中 `sender` 會被寄件者的名稱取代
- `To: receiver` 其中 `receiver` 會被取代為接收者的名稱

# --hints--

你應該使用 f-strings 來列印寄件者的名稱，使用 `From: {self.sender.name}`。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'From: {self.sender.name}')")`));
    }
})
```

你應該使用 f-strings 來列印收件者的名稱，使用 `To: {self.receiver.name}`。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'To: {self.receiver.name}')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
