---
id: 6853f34ff93594934f6636d9
title: 步驟 27
challengeType: 20
dashedName: step-27
---

# --description--

讓我們為你的 `Email` 類別添加字串表現，以便我們可以顯示簡短的電子郵件摘要。

為你的 `Email` 類別添加一個 `__str__` 方法，該方法以 `self` 作為參數。

# --hints--

你應該在 `Email` 類別中定義一個 `__str__` 方法。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__str__")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

--fcc-editable-region--
    
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
