---
id: 6853f4584f508e95cc7a9f4d
title: 第 30 步
challengeType: 20
dashedName: step-30
---

# --description--

現在你將創建一個方法來列出收件匣中的所有電子郵件。當收件匣為空時，這個方法應該能控制代碼，並且當有電子郵件時顯示編號的電子郵件列表。

為你的 `Inbox` 類別添加一個名為 `list_emails` 的方法，該方法以 `self` 作為參數。首先，創建一個 `if` 述語來檢查 `emails` 列表是否為空。如果它是空的，列印訊息 `Your inbox is empty.\n`。添加一個 `return` 述語以離開該方法。

# --hints--

你應該在 `Inbox` 類別中定義一個名為 `list_emails` 的方法，該方法以 `self` 作為參數。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").has_args("self")`))
    }
})
```

你應該使用 `if not self.emails` 來檢查 `self.emails` 是否為空。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }
})
```

如果沒有電子郵件，你應該列印訊息 `Your inbox is empty.\n`。

```js
({
    test: () => 
    { 
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[0].is_equivalent("print('Your inbox is empty.\\\\n')")`))
    }
})
```

你應該在 `if` 述語的結尾有一個 `return` 述語。

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[1].is_equivalent("return")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    
--fcc-editable-region--

```
