---
id: 685941422c30cd3dc41b1d0a
title: 第 22 步
challengeType: 20
dashedName: step-22
---

# --description--

現在讓我們為電子郵件顯示添加標題。在 `display_full_email` 方法中，呼叫 `mark_as_read()` 之後，列印 `\n--- Email ---` 以清楚的標題開始電子郵件顯示。

# --hints--

你應該在 `display_full_email` 方法中列印標題 `\n--- Email ---`。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print('\\\\n--- Email ---')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

```
