---
id: 68594df7779a164e6a61e8ae
title: 步驟 32
challengeType: 20
dashedName: step-32
---

# --description--

收件匣需要一個方法來讀取特定的電子郵件。當使用者想要閱讀電子郵件時，他們會指定想要查看的電子郵件編號（從索引 0 開始），該方法將顯示完整的電子郵件內容。

為你的 `Inbox` 類別添加一個名為 `read_email` 的方法，該方法帶有一個 `index` 參數。首先，檢查收件匣是否為空，如果是，則列印訊息 `Inbox is empty.\n`。之後添加一個 `return` 述語以離開該方法。

# --hints--

你應該在 `Inbox` 類別中定義一個名為 `read_email` 的方法，該方法接受 `self` 和 `index` 作為參數。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").has_args("self, index")`));
    }
})
```

你應該使用 `if not self.emails` 來檢查 `self.emails` 是否為空。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`));
    }
})
```

如果沒有電子郵件，你應該列印 `Inbox is empty.\n` 並傳回。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[0].is_equivalent("print('Inbox is empty.\\\\n')")`));
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[1].is_equivalent("return")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

--fcc-editable-region--
    
--fcc-editable-region--
```
