---
id: 68593ea22eac7e228ed585ce
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

電子郵件應該追蹤是否已被閱讀。為你的 `__init__` 方法添加一個 `read` 屬性，並將其預設值設為 `False`，因為新的電子郵件一開始是未讀取的。

# --hints--

你應該在 `__init__` 方法中將 `False` 指定給 `self.read`。

```js
({ 
    test: () => runPython(`
    assert _Node(_code).find_class("Email").find_function("__init__").find_variable("self.read").is_equivalent("self.read = False")
    `) 
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
```
