---
id: 6859400f8f7f872e50e1238c
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

`User` 類別需要一個方法來傳送電子郵件給其他使用者。當傳送電子郵件時，你會創建一個新的 `Email` 物件並將它添加到接收者的收件匣。

在你的 `User` 類別中創建一個名為 `send_email` 的方法。此方法應該接受 `receiver`、`subject` 和 `body` 的參數。 

在該方法內，使用以下參數值創建一個新的 `Email` 物件，並將其指定給名為 `email` 的變數：

- `sender`：`self`（當前使用者）
- `receiver`：`receiver`（接收電子郵件的使用者）
- `subject`：`subject`（電子郵件的主旨）
- `body`：電子郵件的 `body`（內容）

# --hints--

你應該在 `User` 類別中定義一個 `send_email` 方法。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").has_function("send_email")`))
    }

})
```

`send_email` 方法應該有參數：`self`、`receiver`、`subject` 和 `body`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_args("self, receiver, subject, body")`))
    }

})
```

該方法應該使用正確的參數創建一個名為 `email` 的 `Email` 物件。

```js
({
  test: () => {
    assert(runPython(`
cond = _Node(_code).find_class("User").find_function("send_email").find_body()[0]
cond.is_equivalent("email = Email(sender=self, receiver=receiver, subject=subject, body=body)") or cond.is_equivalent("email = Email(self, receiver, subject, body)")
    `));
  }
});
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
```
