---
id: 6859523c793377706e19e9a0
title: 步驟 42
challengeType: 20
dashedName: step-42
---

# --description--

現在讓我們為電子郵件添加時間戳。在 Email 類別的 `__init__` 方法中，創建一個 `timestamp` 屬性並指定當前時間給它，以自動設定電子郵件建立的時間戳。

這有助於追蹤訊息何時被發送和接收。

# --hints--

你應該在 Email 的 `__init__` 方法中有一個 `self.timestamp`，並將 `datetime.datetime.now()` 指定給它。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").has_stmt("self.timestamp = datetime.datetime.now()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        del self.emails[actual_index]
        print('Email deleted.\n')
```
