---
id: 685c1e32254e5993a5c8823e
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

現在你已經完成了 `Email` 類別的初始設定，是時候創建能夠發送和接收電子郵件的使用者。每個使用者都會有一個名稱，並且能夠執行電子郵件操作。

創建一個 `User` 類別，並包含一個接受 `self` 和 `name` 作為參數的 `__init__` 方法。在類別內，將 `name` 參數指定給 `self.name`。

# --hints--

你應該創建一個 `User` 類別。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_class("User")`))
    }

})
```

`User` 類別應該有一個 `__init__` 方法，該方法接受 `self` 和 `name` 作為參數。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").has_args("self, name")`))
    }

})
```

你應該將 `name` 參數指定給 `self.name`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.name").is_equivalent("self.name = name")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--

--fcc-editable-region--
```
