---
id: 685c1eba966c6396850736ea
title: 步驟 13
challengeType: 20
dashedName: step-13
---

# --description--

雖然使用者可以發送電子郵件並擁有收件匣，但我們需要一個專門的類別來高效管理收件匣的操作。 

`Inbox` 類別將儲存電子郵件列表，並提供添加新電子郵件、列出所有電子郵件及管理它們的方法。

創建一個名為 `Inbox` 的新類別，並包含一個 `__init__` 方法，該方法初始化一個名為 `emails` 的空列表。

# --hints--

你應該創建一個 `Inbox` 類別。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_class("Inbox")`));
    }
})
```

`Inbox` 類別應該有一個接受 `self` 作為參數的 `__init__` 方法。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").has_args("self")`));
    }
})
```

你應該在 `Inbox` 類別的 `__init__` 方法中初始化 `self.emails = []`。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").find_variable("self.emails").is_equivalent("self.emails = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--

--fcc-editable-region--
```
