---
id: 685cdfd1932d8ac6ad986813
title: 步驟 28
challengeType: 20
dashedName: step-28
---

# --description--

在 `__str__` 方法中，你會顯示電子郵件是否已讀取。 

條件運算式允許你在單行中根據條件指定值。

這裡有一個範例：

```py
x = 10
y = 'Even' if x % 2 == 0 else 'Odd' # y will be Even
```

在該方法中，之前，使用條件運算式將字串 `Read` 指定給變數 `status`（如果電子郵件已讀取），如果未讀取則指定為 `Unread`。 

# --hints--

你應該在 `__str__` 方法中有一個 `status` 變數。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__str__").has_variable("status")`));
    }
})
```

你應該將條件表達式指定給 `status` 變數。

```js
({
    test: () => {
        runPython(`
        import ast
        assert isinstance(_Node(_code).find_class("Email").find_function("__str__").find_variable("status").tree.value, ast.IfExp)`);
    }
})
```

你的條件運算式如果電子郵件已讀，應該計算為字串 `Read`，如果未讀，則計算為字串 `Unread`。

```js
({
    test: () => {
        runPython(`
        if_exps = [
          "'Read' if self.read else 'Unread'",
          "'Read' if self.read == True else 'Unread'",
          "'Unread' if not self.read else 'Read'",
          "'Unread' if self.read == False else 'Read'"
        ]
        foo = _Node(_code).find_class("Email").find_function("__str__")
        assert any(foo.has_stmt(f"status = {if_exp}") for if_exp in if_exps)
        `);
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

--fcc-editable-region--
    def __str__(self):
        pass
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
