---
id: 685d2fd1f309b9191f257c55
title: 步驟 40
challengeType: 20
dashedName: step-40
---

# --description--

在將時間戳整合到我們的電子郵件系統之前，讓我們先練習使用日期時間格式化。`datetime.datetime.now()` 函式會給我們目前的日期和時間，我們可以使用 `strftime()` 方法以不同方式格式化它。

以下是 `strftime()` 如何與格式代碼一起運作的方式：

```python
now = datetime.datetime.now()
print(now.strftime("%Y-%m-%d"))  # Output: 2024-03-15 (year-month-day with - separator)
```

格式代碼如 `%Y`（年）、`%m`（月）、`%d`（日）告訴 `strftime()` 要包含什麼，並且你可以在它們之間添加像 `-` 這樣的分隔符。

在你的程式碼底部，創建一個名為 `current_time` 的變數並指定其為 `datetime.datetime.now()`。然後使用 `strftime()` 以 `:` 作為分隔符，列印出小時:分鐘:秒格式的時間。 

使用這些格式代碼：`%H` 代表小時（24 小時制）、`%M` 代表分鐘，`%S` 代表秒數。

# --hints--

你應該創建一個名為 `current_time` 的變數並指定其為 `datetime.datetime.now()`。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_stmt("current_time = datetime.datetime.now()")`));
    }
})
```

你應該使用 `strftime()` 以 `%H:%M:%S` 格式列印時間。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(current_time.strftime('%H:%M:%S'))")`));
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        del self.emails[actual_index]
        print('Email deleted.\n')

--fcc-editable-region--

--fcc-editable-region--
```
