---
id: 688754f3494b17e0b352d67d
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

現在是測試我們完整電子郵件系統的時候了！讓我們創建一些使用者，並觀察電子郵件功能的運作。

創建兩個使用者物件：`alice` 名稱為 `"Alice"`，`bob` 名稱為 `"Bob"`。這將示範使用者如何被創建並相互互動。

# --hints--

你應該創建一個名為 `alice`、名稱為 `"Alice"` 的使用者物件。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("alice").is_equivalent('alice = User("Alice")')`));
    }
})
```

你應該創建一個名為 `bob`、名稱為 `"Bob"` 的使用者物件。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("bob").is_equivalent('bob = User("Bob")')`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--

--fcc-editable-region--
```
