---
id: 6887550fe36f64e27074475a
title: 步驟 18
challengeType: 20
dashedName: step-18
---

# --description--

讓 Alice 寄一封電子郵件給 Bob，以確認電子郵件傳送是否正常運作。

使用 Alice 的 `send_email` 方法將主旨為 `"Hello"`、內容為 `"Hi Bob, how are you?"` 的電子郵件寄給 Bob。

# --hints--

你應該使用正確的參數呼叫 `alice.send_email()`。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call('alice.send_email(bob, "Hello", "Hi Bob, how are you?")')`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
alice = User("Alice")
bob = User("Bob")

--fcc-editable-region--
```
