---
id: 688754d6d59578df1313cb3c
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

現在是時候將 `User` 和 `Inbox` 類別連接起來，讓電子郵件能夠真正傳送！ 

使用者需要擁有適當的 Inbox 物件，而非簡單的列表，且 `send_email` 方法應該將電子郵件傳送到接收者的 inbox。

更新 `User` 類別的 `__init__`，使每個使用者獲得一個實際的 Inbox 物件，而不是先前建立的空列表。 

# --hints--

你應該更新使用者類別以初始化 `self.inbox = Inbox()`。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = Inbox()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []
--fcc-editable-region--

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

class Inbox:
    def __init__(self):
        self.emails = []

```
