---
id: 68875554c6c32de788afc64e
title: 第 21 步
challengeType: 20
dashedName: step-21
---

# --description--

使用者需要一種正確閱讀他們電子郵件的方法。

為你的 `Email` 類別添加一個只帶有 `self` 參數的 `display_full_email` 方法。在此方法內，呼叫 `mark_as_read()` 方法以在有人查看電子郵件時將其標記為已讀。

此方法會用來以美觀的格式顯示電子郵件的完整內容。

# --hints--

你應該在 `Email` 類別中定義一個 `display_full_email` 方法。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").has_function("display_full_email")`));
    }
})
```

`display_full_email` 方法應該只接受 `self` 作為參數。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_args("self")`));
    }
})
```

該方法應該呼叫 `self.mark_as_read()` 來將電子郵件標記為已讀。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("self.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True
--fcc-editable-region--
    
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
