---
id: 68b9366fae8b13c1278761b1
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

你的收件匣需要一種方式來接收新郵件。當有人寄送郵件給使用者時，該郵件應該被添加到他們的收件匣中。

為你的 `Inbox` 類別添加一個名為 `receive_email` 的方法，該方法接受 `self` 和一個電子郵件物件 `email` 作為參數。在方法主體中，使用 `append` 方法將 `email` 添加到 `emails` 列表中。

# --hints--

你應該在 `Inbox` 類別中定義一個名為 `receive_email` 的方法，該方法接受 `self` 和 `email` 作為參數。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_args("self, email")`))
    }
})
```

你應該將 `email` 附加到 `self.emails`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_stmt("self.emails.append(email)")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--
class Inbox:
    def __init__(self):
        self.emails = []

    
--fcc-editable-region--
```
