---
id: 68b930ee0e7b8fb81895e59a
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

現在你將為標記電子郵件為已讀添加一個方法。當有人打開電子郵件時，你想將其狀態從未讀更改為已讀。這個方法很簡單－它只需要將 `read` 屬性設為 `True`。

為你的 `Email` 類別添加一個名為 `mark_as_read` 的方法。此方法應該只接受 `self` 作為參數，並將 `read` 屬性設為 `True`。

# --hints--

你應該在 `Email` 類別中定義一個名為 `mark_as_read` 的方法，該方法只應該接受 `self` 作為參數。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").has_args("self")`))
    }

})
```

該方法應將 `self.read` 設定為 `True`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").find_variable("self.read").is_equivalent("self.read = True")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
```
