---
id: 68b936bf4f47d6c29a669f99
title: 第 16 步
challengeType: 20
dashedName: step-16
---

# --description--

仍在 `send_email` 方法中，呼叫 `receiver` 的 `inbox` 的 `receive_email` 方法，將電子郵件添加到接收者的收件匣。

# --hints--

你應該將 `receiver.inbox.receive_email(email)` 添加到 `send_email` 方法中。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_stmt("receiver.inbox.receive_email(email)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()
        
--fcc-editable-region--
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)
```
