---
id: 68b9350183ff05bd106c79e4
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

現在，測試你在 `Email` 類別中創建的 `mark_as_read` 方法。對現有的 `Email` 物件 `email_obj` 使用該方法以更改狀態。在將電子郵件標記為已讀後，列印 `email_obj` 的 `read` 屬性以確認更改。

# --hints--

你應該對 `email_obj` 呼叫 `mark_as_read` 方法，以將電子郵件標記為已讀。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`))
    }

})
```

你應該在呼叫 `mark_as_read` 方法後列印 `email_obj` 的 `read` 屬性。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).is_ordered("email_obj.mark_as_read()", "print(email_obj.read)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)

--fcc-editable-region--
```
