---
id: 67cd5bace629fa219a70aa6d
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

إذا نظرت إلى وحدة التحكم، قد تتفاجأ برؤية القيمة `false` بدلًا من `true`. يحدث هذا لأن React لا يُحدّث الحالة فورًا. ستُحدّث الحالة فقط في دورة العرض التالية.

خطأ شائع بين المطوّرين الجدد على React هو وضع عبارات وحدة التحكم مباشرة بعد دالة التعيين. لذا يجب أن تكون واعيًا لهذا الأمر أثناء بناء تطبيقات React الخاصة بك.

بما أن `console.log()` لم يعد ضروريًا، احذفه من دالة `handleToggleVisibility` الخاصة بك.

# --hints--

يجب ألا يكون لديك بعد الآن `console.log(isVisible);` في شفرتك.

```js
assert.notMatch(code, /console\.log\s*\(\s*isVisible\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Toggle Visibility</title>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { ToggleApp } from "./index.jsx";
    ReactDOM.createRoot(document.getElementById("root")).render(<ToggleApp />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#toggle-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#toggle-button {
  padding: 10px 20px;
  border: none;
  background: #007BFF;
  color: white;
  border-radius: 5px;
  cursor: pointer;
  transition: background 0.3s ease;
}

#toggle-button:hover {
  background: #0056b3;
}

#message {
  margin-top: 20px;
  font-size: 1.125rem;
  color: #333;
}
```

```jsx
const { useState } = React;

export const ToggleApp = () => {
  const [isVisible, setIsVisible] = useState(false);
  
  const handleToggleVisibility = () => {
    setIsVisible(!isVisible);
  --fcc-editable-region--
    console.log(isVisible);
  --fcc-editable-region--
  }
  
  return (
    <div id="toggle-container">
      <button onClick={handleToggleVisibility} id="toggle-button">
        Message
      </button>
      {isVisible && <p id="message">I love freeCodeCamp!</p>}
    </div>
  );
};
```
