---
id: 5dc174fcf86c76b9248c6eb2
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Los elementos HTML tienen etiquetas de apertura como `<h1>` y etiquetas de cierre como `</h1>`. El texto que un elemento mostrará va entre sus etiquetas de apertura y cierre.

Cambie el texto del elemento `h1` a continuación de `Hello World` a `CatPhotoApp` y observe el cambio en la vista previa del navegador.

Cuando hayas terminado, presiona el botón "Comprueba tu código" para ver si es correcto.

# --hints--

El texto `CatPhotoApp` debe estar presente en el código. Puede que quieras revisar la ortografía.

```js
assert(code.match(/catphotoapp/i));
```

Tu elemento `h1` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Tus etiquetas `h1` deben estar en minúsculas. Por convención, todas las etiquetas HTML se escriben en minúsculas.

```js
assert.notMatch(code, /<\/?H1>/);
```

Tu elemento `h1` debe tener una etiqueta de cierre. Las etiquetas de cierre tiene una `/` después del carácter `<`.

```js
assert(code.match(/<\/h1\>/));
```

Tienes más de un elemento `h1`. Elimina el elemento `h1` extra.

```js
assert(document.querySelectorAll('h1').length === 1);
```

El texto de tu elemento `h1` deber ser `CatPhotoApp`. Probablemente olvidaste añadir el texto, tienes un error tipográfico o no está entre las etiquetas de apertura y cierre del elemento `h1`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

