---
id: 5dc24073f86c76b9248c6ebb
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Los <dfn>atributos</dfn> HTML son palabras especiales usadas dentro de la etiqueta de apertura de un elemento para controlar el comportamiento del elemento. El atributo `src` en un elemento `img` especifica la URL (donde se localiza la imagen).

Aquí hay un ejemplo de un elemento `img` con un atributo `src` apuntando al logo de freeCodeCamp:

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

Dentro del elemento `img` existente, agrega un atributo `src` con esta URL:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

Tu código debe tener un elemento `img`. Pudiste haber borrado el elemento `img` o no pusiste entre comillas el valor del atributo `src`.

```js
assert(document.querySelector('img'));
```

Tu elemento `img` debe tener un atributo `src`. Probablemente no has añadido el atributo o tienes un error tipográfico. Asegúrate e de que hay un espacio entre el nombre del elemento y el nombre del atributo.

```js
assert(document.querySelector('img').src);
```

El atributo `src` del elemento `img` debe tener el valor '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'. Probablemente no has añadido la URL o tienes un error tipográfico. La URL debe ser exactamente igual, cuidando mayúsculas y minúsculas.

```js
assert.strictEqual(document.querySelector('img').src, 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

Aunque le has dado la URL correcta al parámetro `src` del elemento `img`, se recomienda siempre poner el valor de un atributo entre comillas.

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

Tu elemento `img` debe cerrarse ya sea con `>` o `/>`.

```js
assert.match(code, /<img\s+src\s*=\s*("|')\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\s*\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```

