---
id: 5dfb6a35eacea3f48c6300b4
title: Paso 26
challengeType: 0
dashedName: step-26
---

# --description--

Un elemento de subtítulo de figura (`figcaption`) se usa para agregar una leyenda que describe la imagen contenida dentro del elemento `figure`.

Aquí tienes un ejemplo de un elemento `figcaption` con la leyenda de `A cute cat`:

```html
<figure>
  <img src="image.jpg" alt="A description of the image">
  <figcaption>A cute cat</figcaption>
</figure>
```

Después de la imagen anidada en el elemento `figure`, agrega un elemento `figcaption` con el texto fijado en:

`Cats love lasagna.`

# --hints--

El elemento `figcaption` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen la siguiente sintaxis: `<elementName>`.

```js
assert(document.querySelector('figcaption'));
```

El elemento `figcaption` debe tener una etiqueta de cierre. Las etiquetas de cierre tienen un `/` justo después del carácter `<`.

```js
assert(code.match(/<\/figcaption\>/));
```

El elemento `figcaption` debe estar anidado en el elemento `figure`.

```js
assert(
  document.querySelector('figure > figcaption') &&
    document.querySelector('figure > figcaption')
);
```

El elemento `img` de lasaña debe estar anidado en el elemento `figure`.

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

El elemento `figcaption` anidado en el elemento `figure` debe estar debajo del elemento `img`. Los tienes en el orden incorrecto.

```js
assert(
  document.querySelector('figcaption').previousElementSibling.nodeName === 'IMG'
);
```

El texto de tu elemento `figcaption` debe ser `Cats love lasagna.` Has omitido el texto o tienes un error tipográfico.

```js
assert(
  document.querySelector('figcaption').innerText.trim().match(/^Cats love lasagna\.?$/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

