---
id: 5ef9b03c81a63668521804dc
title: Paso 46
challengeType: 0
dashedName: step-46
---

# --description--

Puedes usar botones de radio (radio button) para preguntas en las que quieres que el usuario solo te dé una respuesta a partir de múltiples opciones.

Aquí hay un ejemplo de un botón de radio con el texto establecido como `cat`:

```html
<input type="radio"> cat
```

Recuerda que un elemento de `input` es un elemento vacío.

Antes de la entrada de texto, añada un botón de radio con el texto establecido como:

`Indoor`

# --hints--

Debes crear un elemento `input` para tu botón de radio. Revisa la sintaxis.

```js
assert(document.querySelectorAll('form > input').length >= 2);
```

Tu elemento `input` debe tener una etiqueta de apertura, pero no una etiquete de cierre.

```js
assert(document.querySelector('form > input') && !code.match(/<\/input\>/g));
```

Debes tener solo un elemento `input` agregado para tu botón de radio. Elimine cualquier extra.

```js
assert(document.querySelectorAll('form > input'), 2);
```

Tu nuevo elemento `input` debe estar arriba del elemento `input` ya existente que tiene el atributo `type` con el valor `text`. Los tienes en el orden incorrecto.

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

Tu nuevo elemento `input` no tiene un atributo `type`. Comprueba que hay un espacio después del nombre de la etiqueta de apertura.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

Tu nuevo elemento `input` debe tener solo un atributo `type`. Elimine cualquier extra.

```js
assert(document.querySelector('input')
  ?.getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

Tu nuevo elemento `input` debe tener un atributo `type` con el valor `radio`. Probablemente no has añadido el valor o tienes un error tipográfico. Recuerda que los valores de los atributos deben estar entre comillas.

```js
assert.match(
  document.querySelector('input')?.getAttribute('type'),
  /^radio$/i
);
```

Aunque al atributo `type` del elemento `input` le hayas dado el valor `radio`, se recomienda siempre poner el valor de un atributo entre comillas.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

El texto `Indoor` del botón `radio` debe estar ubicado después de este en lugar de antes.

```js
const radioInputElem = document.querySelector('input');
assert.notMatch(radioInputElem?.previousSibling?.nodeValue, /Indoor/i);
```

El texto ` Indoor` debe estar ubicado directamente a la derecha de tu botón `radio`. Has omitido el texto o tienes un error tipográfico.

```js
const radioInputElem = document.querySelector('input');
assert.match(radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '), /\s*Indoor/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
 
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

