---
id: 5ef9b03c81a63668521804d2
title: Paso 29
challengeType: 0
dashedName: step-29
---

# --description--

El código de una lista ordenada, ordered list (`ol`), es similar al de una lista no ordenada, unordered list (ul), pero los elemento de una lista ordenada aparecen enumerados.

Debajo del elemento `h3`, agrega una lista ordenada con estos tres elementos de lista:

`flea treatment`
`thunder`
`other cats`

# --hints--

Tu elemento `ol` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert(document.querySelector('ol'));
```

Tu elemento `ol` debe tener una etiqueta de cierre. Las etiquetas de cierre tiene una `/` después del carácter `<`.

```js
assert(code.match(/<\/ol>/));
```

El elemento `ol` debe estar arriba de la etiqueta de cierre del segundo elemento `section`. Los tienes en el orden incorrecto.

```js
assert.equal(document.querySelectorAll('main > section')?.[1]?.lastElementChild?.nodeName, 'OL');
```

Los elementos `li` deben estar anidados dentro del elemento `ol`.

```js
assert(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ).length === 3
);
```

Debes tener tres elementos `li` con el texto`flea treatment`, `thunder` y`other cats` en cualquier orden.

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText.toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

Debes tener solo un elemento `ol`.

```js 
assert([...document.querySelectorAll('ol')].length == 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--
        <h3>Top 3 things cats hate:</h3>
        
--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

