---
id: 5a24c314108439a4d403615a
title: Ondoa kipengee kutoka kwenye safu ya data
challengeType: 6
forumTopicId: 301447
dashedName: remove-an-item-from-an-array
---

# --description--

Ni wakati wa kufanya mazoezi ya kuondoa vipengee kutoka kwenye safu ya data. Opereta ya spread inaweza kutumika pia hapa. Njia nyingine muhimu za JavaScript ni `slice()` na `concat()`.

# --instructions--

Reducer na muundaji wa kitendo vimebadilishwa kuondoa kipengee kutoka kwenye safu ya data kulingana na kielezo cha kipengee hicho. Malizia kuandika reducer ili kurejesha safu mpya ya hali yenye kipengee kilicho kwenye kielezo maalum kimeondolewa.

# --hints--

Hifadhi ya Redux inapaswa kuwepo na kuanzishwa na hali sawa na `[0,1,2,3,4,5]`

```js
const initialState = store.getState();
assert.isArray(initialState) 
assert.deepEqual(initialState, [0, 1, 2, 3, 4, 5])
```

`removeItem` na `immutableReducer` zote zinapaswa kuwa vitendakazi.

```js
assert(
  typeof removeItem === 'function' && typeof immutableReducer === 'function'
);
```

Kutuma kitendo cha muundaji wa `removeItem` kunapaswa kuondoa vipengee kutoka kwenye hali na HAKUTAKI kubadilisha hali hiyo moja kwa moja.

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(removeItem(3));
const state_1 = store.getState();
store.dispatch(removeItem(2));
const state_2 = store.getState();
store.dispatch(removeItem(0));
store.dispatch(removeItem(0));
store.dispatch(removeItem(0));
const state_3 = store.getState();

assert(isFrozen)
assert.deepEqual(state_1, [0, 1, 2, 4, 5]) 
assert.deepEqual(state_2, [0, 1, 4, 5]) 
assert.deepEqual(state_3, [5])

```

# --seed--

## --seed-contents--

```js
const immutableReducer = (state = [0,1,2,3,4,5], action) => {
  switch(action.type) {
    case 'REMOVE_ITEM':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const removeItem = (index) => {
  return {
    type: 'REMOVE_ITEM',
    index
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const immutableReducer = (state = [0,1,2,3,4,5], action) => {
  switch(action.type) {
    case 'REMOVE_ITEM':
      return [
        ...state.slice(0, action.index),
        ...state.slice(action.index + 1)
      ];
    default:
      return state;
  }
};

const removeItem = (index) => {
  return {
    type: 'REMOVE_ITEM',
    index
  }
}

const store = Redux.createStore(immutableReducer);
```
