---
id: 5a24c314108439a4d403614c
title: "Pata state kutoka hifadhi ya Redux"
challengeType: 6
forumTopicId: 301443
dashedName: get-state-from-the-redux-store
---

# --description--

Kipengele cha hifadhi ya Redux kinatoa njia kadhaa zinazokuwezesha kuingiliana nacho. Kwa mfano, unaweza kupata `state` ya sasa iliyoko katika kipengele cha hifadhi ya Redux kwa kutumia njia ya `getState()`.

# --instructions--

Msimbo kutoka changamoto iliyopita umeandikwa upya kwa ufupi zaidi katika mhariri wa msimbo. Tumia `store.getState()` kupata `state` kutoka `store`, na weka thamani hii katika kigezo kipya `currentState`.

# --hints--

Hifadhi ya Redux inapaswa kuwa na thamani ya 5 kwa state ya awali.

```js
assert(store.getState() === 5);
```

Kigezo `currentState` kinapaswa kuwepo na kimepewa thamani ya state ya sasa ya hifadhi ya Redux.

```js
assert(currentState === 5 && code.includes('store.getState()'));
```

# --seed--

## --seed-contents--

```js
const store = Redux.createStore(
  (state = 5) => state
);

// Change code below this line
```

# --solutions--

```js
const store = Redux.createStore(
  (state = 5) => state
);

// Change code below this line
const currentState = store.getState();
```
