---
id: 5a24c314108439a4d4036158
title: Usibadilishe hali moja kwa moja
challengeType: 6
forumTopicId: 301445
dashedName: never-mutate-state
---

# --description--

Changamoto hizi za mwisho zinaelezea njia kadhaa za kutekeleza kanuni kuu ya kutokubadilisha hali moja moja katika Redux. Hali isiyobadilika ina maana kwamba hauwezi kubadilisha hali moja moja, badala yake, unarudisha nakala mpya ya hali hiyo.

Ikiwa unachukua picha ya hali ya programu ya Redux kwa muda, utaona kitu kama `state 1`, `state 2`, `state 3`,`state 4`, `...` na kadhalika ambapo kila hali inaweza kuwa sawa na ile ya mwisho, lakini kila moja ni kipande tofauti cha data. Kutokubadilika kwa hali hii, kwa kweli, ndiko kunakotoa vipengele kama kurekebisha makosa kwa kutumia kusafiri kwa wakati ambacho huenda umesikia kuhusu.

Redux haitekelezi kwa nguvu kutokubadilika kwa hali katika hifadhi yake au reducers, jukumu hilo ni la mpogramu. Kwa bahati nzuri, JavaScript (hasa ES6) hutoa zana kadhaa muhimu unazoweza kutumia kutekeleza kutokubadilika kwa hali yako, iwe ni `string`, `number`, `array`, au `object`. Kumbuka kwamba mfuatano wa herufi na nambari ni thamani za msingi na hazibadiliki kwa asili. Kwa maneno mengine, 3 daima ni 3. Huwezi kubadilisha thamani ya nambari 3. Hata hivyo, `array` au `object` ni zinazobadilika. Kwa vitendo, hali yako huenda ikajumuisha `array` au `object`, kwani hizi ni miundo ya data muhimu kwa kuwakilisha aina nyingi za taarifa.

# --instructions--

Kuna `store` na `reducer` katika mhariri wa msimbo kwa ajili ya kusimamia vitu vya kufanya. Malizia kuandika kesi ya `ADD_TO_DO` katika reducer ili kuongeza kipengele kipya cha kufanya kwenye hali. Kuna njia kadhaa za kufanya hivi kwa kutumia JavaScript ya kawaida au ES6. Angalia kama unaweza kupata njia ya kurudisha safu mpya ya data yenye kipengele kutoka `action.todo` kimeongezwa mwishoni.

# --hints--

Hifadhi ya Redux inapaswa kuwepo na kuanzishwa na hali sawa na safu ya `todos` katika mhariri wa msimbo.

```js
assert(
  (function () {
    const todos = [
      'Go to the store',
      'Clean the house',
      'Cook dinner',
      'Learn to code'
    ];
    const initialState = store.getState();
    return (
      Array.isArray(initialState) && initialState.join(',') === todos.join(',')
    );
  })()
);
```

`addToDo` na `immutableReducer` zote zinapaswa kuwa njia.

```js
assert(typeof addToDo === 'function' && typeof immutableReducer === 'function');
```

Kutuma kitendo cha aina `ADD_TO_DO` kwenye hifadhi ya Redux kunapaswa kuongeza kipengele cha `todo` na HAKUTAKI kubadilisha hali moja moja.

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(addToDo('__TEST__TO__DO__'));
const finalState = store.getState();
const expectedState = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
  '__TEST__TO__DO__'
];
assert(isFrozen);
assert.deepEqual(finalState, expectedState);
```

# --seed--

## --seed-contents--

```js
const ADD_TO_DO = 'ADD_TO_DO';

// A list of strings representing tasks to do:
const todos = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
];

const immutableReducer = (state = todos, action) => {
  switch(action.type) {
    case ADD_TO_DO:
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: ADD_TO_DO,
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const ADD_TO_DO = 'ADD_TO_DO';

const todos = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
];

const immutableReducer = (state = todos, action) => {
  switch(action.type) {
    case ADD_TO_DO:
      return state.concat(action.todo);
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: ADD_TO_DO,
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```
