---
id: 5a24c314108439a4d403614d
title: Tangaza tukio la Redux
challengeType: 6
forumTopicId: 301440
dashedName: define-a-redux-action
---

# --description--

Kwa kuwa Redux ni mfumo wa kazi wa usimamizi wa hali, kusahihisha hali ni mojawapo ya kazi zake kuu. Katika Redux, masasisho yote ya hali huanzishwa kwa kutuma matukio. Tukio ni tu kitu cha JavaScript chenye taarifa kuhusu tukio lililotokea. Hifadhi ya Redux hupokea vitu hivi vya tukio, kisha husahihisha hali yake ipasavyo. Wakati mwingine tukio la Redux pia hubeba data fulani. Kwa mfano, tukio hubeba jina la mtumiaji baada ya mtumiaji kuingia. Ingawa data ni hiari, matukio lazima yabe na sifa `type` inayobainisha aina ya tukio lililotokea.

Fikiria matukio ya Redux kama wajumbe wanaotoa taarifa kuhusu matukio yanayotokea katika programu yako kwa hifadhi ya Redux. Hifadhi hiyo kisha hufanya kazi ya kusahihisha hali kulingana na tukio lililotokea.

# --instructions--

Kuandika tukio la Redux ni rahisi kama kutangaza kitu chenye sifa ya aina. Tangaza kitu `action` na kipa sifa `type` iliyowekwa kwa mfuatano wa herufi `'LOGIN'`.

# --hints--

Kitu cha `action` kinapaswa kuwepo.

```js
assert(
  (function () {
    return typeof action === 'object';
  })()
);
```

Kitu cha `action` kinapaswa kuwa na sifa ya funguo `type` yenye thamani `LOGIN`.

```js
assert(
  (function () {
    return action.type === 'LOGIN';
  })()
);
```

# --seed--

## --seed-contents--

```js
// Define an action here:
```

# --solutions--

```js
const action = {
  type: 'LOGIN'
}
```
