---
id: 5a24c314108439a4d4036152
title: Tumia const kwa aina za kitendo
challengeType: 6
forumTopicId: 301450
dashedName: use-const-for-action-types
---

# --description--

Mkataba wa kawaida unapotumia Redux ni kuweka aina za kitendo kama constants zisizobadilika, kisha kurejea constants hizi kila mahali zinapotumika. Unaweza kuboresha msimbo unaofanya kazi nao kwa kuandika aina za kitendo kama tamko za `const`.

# --instructions--

Tangaza `LOGIN` na `LOGOUT` kama thamani za `const` na uwaweke kwenye mfuatano wa herufi `'LOGIN'` na `'LOGOUT'`, mtawalia. Kisha, hariri `authReducer()` na watengeneza vitendo kurejea constants hizi badala ya thamani za mfuatano wa herufi.

**Kumbuka:** Kwa kawaida ni mkataba kuandika constants kwa herufi kubwa zote, na hii ni desturi ya kawaida katika Redux pia.

# --hints--

Kuitisha kitendakazi `loginUser` inapaswa kurudisha kitu chenye sifa ya `type` ikiwa na thamani ya mfuatano wa herufi `LOGIN`.

```js
assert(loginUser().type === 'LOGIN');
```

Kuitisha kitendakazi `logoutUser` inapaswa kurudisha kitu chenye sifa ya `type` ikiwa na thamani ya mfuatano wa herufi `LOGOUT`.

```js
assert(logoutUser().type === 'LOGOUT');
```

Hifadhi inapaswa kuanzishwa na kitu chenye sifa ya `login` ikiwa na thamani ya `false`.

```js
assert(store.getState().authenticated === false);
```

Kutuma `loginUser` kunapaswa kusasisha sifa ya `login` katika hali ya hifadhi kuwa `true`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(loginUser());
    const afterLogin = store.getState();
    return (
      initialState.authenticated === false && afterLogin.authenticated === true
    );
  })()
);
```

Kutuma `logoutUser` kunapaswa kusasisha sifa ya `login` katika hali ya hifadhi kuwa `false`.

```js
assert(
  (function () {
    store.dispatch(loginUser());
    const loggedIn = store.getState();
    store.dispatch(logoutUser());
    const afterLogout = store.getState();
    return (
      loggedIn.authenticated === true && afterLogout.authenticated === false
    );
  })()
);
```

Kitendakazi cha `authReducer` kinapaswa kushughulikia aina nyingi za vitendo kwa kutumia taarifa ya kubadilisha.

```js
assert(
    (function () {
      return (
        typeof authReducer === 'function' &&
        code.toString().includes('switch') &&
        code.toString().includes('case') &&
        code.toString().includes('default')
      );
    })()
  );
```

`LOGIN` na `LOGOUT` zinapaswa kutangazwa kama thamani za `const` na kuwekwa mfuatano wa herufi wa `LOGIN` na `LOGOUT`.

```js
const noWhiteSpace = __helpers.removeWhiteSpace(code);
assert(LOGIN === 'LOGIN' && LOGOUT === 'LOGOUT')
assert(noWhiteSpace.includes('const'))
```

Watengeneza vitendo na reducer wanapaswa kurejea constants za `LOGIN` na `LOGOUT`.

```js
  assert(
    (function () {
      const noWhiteSpace = __helpers.removeWhiteSpace(
        code.toString()
      );
      return (
        noWhiteSpace.includes('caseLOGIN:') &&
        noWhiteSpace.includes('caseLOGOUT:') &&
        noWhiteSpace.includes('type:LOGIN') &&
        noWhiteSpace.includes('type:LOGOUT')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```js


const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {
    case 'LOGIN': 
      return {
        authenticated: true
      }
    case 'LOGOUT': 
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```

# --solutions--

```js
const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {

    case LOGIN:
      return {
        authenticated: true
      }

    case LOGOUT:
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: LOGIN
  }
};

const logoutUser = () => {
  return {
    type: LOGOUT
  }
};
```
