---
id: 5a24c314108439a4d4036154
title: Unganisha Reducers Nyingi
challengeType: 6
forumTopicId: 301436
dashedName: combine-multiple-reducers
---

# --description--

Wakati hali ya programu yako inaanza kuwa ngumu zaidi, inaweza kuvutia kugawanya hali katika vipande vingi. Badala yake, kumbuka kanuni ya kwanza ya Redux: hali yote ya programu huhifadhiwa katika kitu kimoja cha hali katika hifadhi. Kwa hiyo, Redux hutoa muundo wa reducers kama suluhisho kwa mfano wa hali tata. Unaeleza reducers nyingi kushughulikia vipande tofauti vya hali ya programu yako, kisha unazunganisha reducers hizi pamoja kuwa reducer kuu moja. Reducer kuu kisha hupitishwa katika njia ya Redux `createStore()`.

Ili kutuwezesha kuunganisha reducers nyingi pamoja, Redux hutoa njia ya `combineReducers()`. Njia hii inakubali kitu kama hoja ambacho unaeleza sifa zinazohusisha funguo na kitendakazi maalum cha reducer. Jina unalolipa funguo litumike na Redux kama jina la kipande cha hali kinachohusiana.

Kawaida, ni desturi nzuri kuunda reducer kwa kila kipande cha hali ya programu wakati ni tofauti au ya kipekee kwa namna fulani. Kwa mfano, katika programu ya kuandika maelezo na uthibitisho wa mtumizi, reducer moja inaweza kushughulikia uthibitisho wa utambulisho wakati mwingine hushughulikia maandishi na maelezo ambayo mtumizi anawasilisha. Kwa programu kama hiyo, tunaweza kuandika njia ya `combineReducers()` kama ifuatavyo:

```js
const rootReducer = Redux.combineReducers({
  auth: authenticationReducer,
  notes: notesReducer
});
```

Sasa, funguo `notes` itakuwa na hali yote inayohusiana na maelezo yetu na inayoshughulikiwa na `notesReducer`. Hivi ndivyo reducers nyingi zinavyoweza kuunganishwa kusimamia hali tata ya programu. Katika mfano huu, hali iliyohifadhiwa katika hifadhi ya Redux itakuwa kitu kimoja chenye sifa za `auth` na `notes`.

# --instructions--

Kuna kazi za `counterReducer()` na `authReducer()` zilizotolewa katika mhariri wa msimbo, pamoja na hifadhi ya Redux. Malizia kuandika kitendakazi cha `rootReducer()` kwa kutumia njia ya `Redux.combineReducers()`. Weka `counterReducer` kwenye funguo iitwayo `count` na `authReducer` kwenye funguo iitwayo `auth`.

# --hints--

`counterReducer` inapaswa kuongeza na kupunguza `state`.

```js
assert(
  (function () {
    const initialState = store.getState().count;
    store.dispatch({ type: INCREMENT });
    store.dispatch({ type: INCREMENT });
    const firstState = store.getState().count;
    store.dispatch({ type: DECREMENT });
    const secondState = store.getState().count;
    return firstState === initialState + 2 && secondState === firstState - 1;
  })()
);
```

`authReducer` inapaswa kubadilisha `state` ya `authenticated` kati ya `true` na `false`.

```js
assert(
  (function () {
    store.dispatch({ type: LOGIN });
    const loggedIn = store.getState().auth.authenticated;
    store.dispatch({ type: LOGOUT });
    const loggedOut = store.getState().auth.authenticated;
    return loggedIn === true && loggedOut === false;
  })()
);
```

Hifadhi ya `state` inapaswa kuwa na funguo mbili: `count`, inayoshikilia nambari, na `auth`, inayoshikilia kitu. Kitu cha `auth` kinapaswa kuwa na sifa ya `authenticated`, inayoshikilia thamani ya kweli au siyo kweli.

```js
assert(
  (function () {
    const state = store.getState();
    return (
      typeof state.auth === 'object' &&
      typeof state.auth.authenticated === 'boolean' &&
      typeof state.count === 'number'
    );
  })()
);
```

`rootReducer` inapaswa kuwa kitendakazi kinachounganisha `counterReducer` na `authReducer`.

```js
assert(
    (function () {
      const noWhiteSpace = __helpers.removeWhiteSpace(code);
      return (
        typeof rootReducer === 'function' &&
        noWhiteSpace.includes('Redux.combineReducers')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const authReducer = (state = {authenticated: false}, action) => {
  switch(action.type) {
    case LOGIN:
      return {
        authenticated: true
      }
    case LOGOUT:
      return {
        authenticated: false
      }
    default:
      return state;
  }
};

const rootReducer = // Define the root reducer here

const store = Redux.createStore(rootReducer);
```

# --solutions--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const authReducer = (state = {authenticated: false}, action) => {
  switch(action.type) {
    case LOGIN:
      return {
        authenticated: true
      }
    case LOGOUT:
      return {
        authenticated: false
      }
    default:
      return state;
  }
};

const rootReducer = Redux.combineReducers({
  count: counterReducer,
  auth: authReducer
});

const store = Redux.createStore(rootReducer);
```
