---
id: 5a24c314108439a4d403614e
title: Tafsiri kitendakazi cha kitendo
challengeType: 6
forumTopicId: 301441
dashedName: define-an-action-creator
---

# --description--

Baada ya kuunda kitendo, hatua inayofuata ni kutuma kitendo hicho kwa hifadhi ya Redux ili iweze kusasisha hali yake. Katika Redux, unaeleza watengenezaji wa vitendo ili kufanikisha hili. Mtengenezaji wa kitendo ni kitendakazi cha JavaScript kinachorudisha kitendo. Kwa maneno mengine, watengenezaji wa vitendo huunda vitu vinavyowakilisha matukio ya kitendo.

# --instructions--

Tafsiri kitendakazi kinachoitwa `actionCreator()` kinachorudisha kitu cha `action` kinapoitwa.

# --hints--

Kitendakazi `actionCreator` kinapaswa kuwepo.

```js
assert(typeof actionCreator === 'function');
```

Kukimbia kitendakazi cha `actionCreator` kunapaswa kurudisha kitu cha `action`.

```js
assert(typeof action === 'object');
```

Kitu kinachorudishwa cha `action` kinapaswa kuwa na sifa ya funguo `type` yenye thamani `LOGIN`.

```js
assert(action.type === 'LOGIN');
```

# --seed--

## --seed-contents--

```js
const action = {
  type: 'LOGIN'
}
// Define an action creator here:
```

# --solutions--

```js
const action = {
  type: 'LOGIN'
}
const actionCreator = () => {
  return action;
};
```
