---
id: 5a24c314108439a4d4036153
title: Sajili msikilizaji wa hifadhi
challengeType: 6
forumTopicId: 301446
dashedName: register-a-store-listener
---

# --description--

Njia nyingine unayopata upatikanaji nayo kwenye kitu cha Redux `store` ni `store.subscribe()`. Hii inakuwezesha kujiandikisha kwa njia za kusikiliza kwenye hifadhi, ambazo huitwa kila wakati kitendo kinapotumwa dhidi ya hifadhi. Matumizi rahisi ya njia hii ni kujiandikisha kitendakazi kwenye hifadhi yako ambacho kinarekodi ujumbe kila wakati kitendo kinapopokelewa na hifadhi inaposasishwa.

# --instructions--

Andika kitendakazi cha mwitiko kinachoongeza thamani ya kigezo cha kimataifa `count` kila wakati hifadhi inapopokea kitendo, na pita kitendakazi hiki kwenye njia ya `store.subscribe()`. Utaona kuwa `store.dispatch()` inaitwa mara tatu mfululizo, kila mara ikipitia moja kwa moja kitu cha kitendo. Tazama matokeo ya konsoli kati ya kutumwa kwa vitendo ili kuona masasisho yanavyotokea.

# --hints--

Kutuma kitendo cha `ADD` kwenye hifadhi kunapaswa kuongeza hali kwa `1`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch({ type: 'ADD' });
    const newState = store.getState();
    return newState === initialState + 1;
  })()
);
```

Kufaa kuwe na kitendakazi cha msikilizaji kilichojiandikisha kwenye hifadhi kwa kutumia `store.subscribe`.

```js
assert.match(code, /store\s*\.\s*subscribe\(/gm);
```

`store.subscribe` inapaswa kupokea kitendakazi.

```js
assert.match(code, /(\s*function\s*)|(\s*\(\s*\)\s*=>)/gm);
```

Kitendakazi kilichopitishwa kwa `store.subscribe` hakipaswi kuitwa.

```js
assert.notMatch(code, /store\.subscribe\(.+\(\)\)/);
```

Wito wa kurudisha kwa `store.subscribe` pia unapaswa kuongeza kigezo cha kimataifa `count` wakati hifadhi inaposasishwa.

```js
assert.strictEqual(store.getState(), count);
```

# --seed--

## --seed-contents--

```js
const ADD = 'ADD';

const reducer = (state = 0, action) => {
  switch(action.type) {
    case ADD:
      return state + 1;
    default:
      return state;
  }
};

const store = Redux.createStore(reducer);

// Global count variable:
let count = 0;

// Change code below this line

// Change code above this line

store.dispatch({type: ADD});
console.log(count);
store.dispatch({type: ADD});
console.log(count);
store.dispatch({type: ADD});
console.log(count);
```

# --solutions--

```js
const ADD = 'ADD';

const reducer = (state = 0, action) => {
  switch(action.type) {
    case ADD:
      return state + 1;
    default:
      return state;
  }
};

const store = Redux.createStore(reducer);
 let count = 0;
// Change code below this line

store.subscribe( () =>
 {
 count++;
 }
);

// Change code above this line

store.dispatch({type: ADD});
store.dispatch({type: ADD});
store.dispatch({type: ADD});
```
