---
id: 662a6bc12cde72c32fb526f0
title: خطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

الـ<dfn>واجهة</dfn> تشبه المخطط الأساسي للصنف. تحتوي الواجهة على مجموعة من الطرق والخصائص التي يجب على الصنف تنفيذها.

ابدأ هذا المشروع بإعلان صنف فارغ باسم `Equation`. ستستخدم هذا الصنف لتعريف واجهة، أي مخطط أساسي لمعادلة عامة.

# --hints--

يجب تعريف صنف جديد باسم `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
