---
id: 662bd552e1c1d2db1b88ba47
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

الآن، عرّف صنفًا آخر باسم `LinearEquation` واجعله يرث من `Equation`. ستستخدم هذا الصنف لتمثيل المعادلات الخطية.

# --hints--

يجب أن تعرّف صنفًا باسم `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("LinearEquation")`)) })
```

يجب أن يرث صنف `LinearEquation` من صنف `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").inherits_from("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
--fcc-editable-region--

--fcc-editable-region--
```
