---
id: 662bd456896f16d9bd03f1a6
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

داخل صنف `Equation`، عرّف طريقتين جديدتين للحالة باسمَي `solve` و `analyze`.

# --hints--

يجب أن تعرّف طريقة باسم `solve` داخل صنف `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("solve")`)) })
```

يجب أن تأخذ طريقتك `solve` معلمة واحدة، وهي `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_args("self")`)) })
```

يجب أن تعرّف طريقة باسم `analyze` داخل صنف `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("analyze")`)) })
```

يجب أن تأخذ طريقتك `analyze` معلمة واحدة، وهي `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Equation:
    pass
--fcc-editable-region--
```
