---
id: 662bdd364bf2cde1487922a9
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

على عكس لغات البرمجة الأخرى، لا ينفذ Python الواجهات في لغته الأساسية، لكن مكتبة Python القياسية تسمح لك بتعريف الواجهات بطريقة بسيطة.

لهذا المشروع، ستستخدم أدوات من وحدة `abc`. لذلك، استورد هذه الوحدة في شفرتك.

# --hints--

يجب استيراد وحدة `abc`.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("import abc")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
    

eq = Equation()
lin_eq = LinearEquation()
```
