---
id: 662bd8260da84bdd5feae419
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

تريد أن يقوم صنف `LinearEquation` بتنفيذ وليس مجرد وراثة كل الطرق المعرفة داخل صنف `Equation`، الذي يجب أن يعمل كواجهة.

حاليًا، صنف `Equation` هو ببساطة الصنف الأصل لـ `LinearEquation`. في الخطوات التالية ستتعلم كيف تحوله إلى واجهة رسمية.

في الوقت الحالي، أنشئ مثيلًا من `Equation` وعيّنه إلى متغير `eq`، ومثيلًا من `LinearEquation` وعيّنه إلى متغير `lin_eq`.

# --hints--

يجب إعلان متغير `eq` وتعيين مثيل من `Equation` له.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("eq = Equation()")`)) })
```

يجب إعلان متغير `lin_eq` وتعيين مثيل من `LinearEquation` له.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation()")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
--fcc-editable-region--

--fcc-editable-region--
```
