---
id: 662bde88dc84f1e249801b1a
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

`ABC` تعني *الأصناف الأساسية المجردة*. يتيح لك صنف `ABC` تحويل صنف عادي إلى صنف مجرد، وهو صنف يعمل كخطة أساسية للأصناف الملموسة.

عدّل عبارة `import` الخاصة بك لاستيراد صنف `ABC` فقط من وحدة `abc`. يمكنك استيراد كائن محدد `x` من وحدة `y` باتباع تركيب الاستيراد `from y import x`.

بعد ذلك، حوّل صنف `Equation` إلى صنف مجرد بجعله يرث من `ABC`.

# --hints--

يجب أن تستورد `ABC` من وحدة `abc`.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("from abc import ABC")`)) })
```

يجب أن يرث صنف `Equation` من `ABC`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").inherits_from("ABC")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import abc


class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass
    
eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
