---
id: 662f96576ef178927de87975
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

لكي تُعترف الطريقة كطريقة مجردة، يجب تزيينها بالزخرفة `@abstractmethod`.

يُستخدم <dfn>الزخرف</dfn> في بايثون لتعديل سلوك الدالة. إليك مثالًا على كيفية استخدام زخرف باسم `spam`:

```py
@spam
def foo():
    pass
```

عدّل جملة الاستيراد لاستيراد زخرف `abstractmethod` وزيّن كلًا من الطريقتين `solve` و`analyze` في صنف `Equation`. هذا سيؤدي إلى ظهور استثنائين.

بمجرد أن يحتوي صنف يرث من `ABC` على طريقة مجردة، لا يمكن إنشاء مثيل من هذا الصنف بعد ذلك. لذلك، احذف مثيل `Equation` للتخلص من الخطأ.

الخطأ الآخر يحدث لأن صنف `LinearEquation` يجب أن ينفذ كل الطرق المجردة المعرفة في الواجهة. تأكد من تعريفها داخل صنف `LinearEquation` أيضًا. لا يجب استخدام زخرف `abstractmethod` في الصنف الملموس.

# --hints--

يجب استيراد `abstractmethod` من وحدة `abc`.

```js
({ test: () => assert(runPython(`
_Node(_code).has_import("from abc import ABC, abstractmethod") or \\
_Node(_code).has_import("from abc import abstractmethod, ABC") or \\
(_Node(_code).has_import("from abc import abstractmethod") and _Node(_code).has_import("from abc import ABC"))
`)) })
```

يجب تزيين طريقة `solve` داخل صنف `Equation` بالزخرف `@abstractmethod`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_decorators("abstractmethod")`)) })
```

يجب تزيين طريقة `analyze` داخل صنف `Equation` بالزخرف `@abstractmethod`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_decorators("abstractmethod")`)) })
```

يجب تعريف طريقة باسم `solve` داخل صنف `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("solve")`)) })
```

يجب أن تأخذ طريقة `solve` معلمة واحدة، وهي `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_args("self")`)) })
```

يجب تعريف طريقة باسم `analyze` داخل صنف `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("analyze")`)) })
```

يجب أن تأخذ طريقة `solve` معلمة واحدة، وهي `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC


class Equation(ABC):
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass

eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
