---
id: 662f6d7c92381a3049e4c987
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

لا يجب أن يعرّف الواجهة طرقًا مجردة فقط، بل يمكنه أيضًا تنفيذ طرق تُورّثها الأصناف الملموسة.

قبل تنفيذ `solve` و`analyze` فعليًا، عرّف طريقة `__init__` داخل صنف `Equation`. لا تستخدم أي مزخرف عليها.

# --hints--

يجب أن تعرّف طريقة `__init__` في صنف `Equation` الخاص بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__init__")`)) })
```

يجب أن تأخذ طريقة `__init__` معلمة واحدة، وهي `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC, abstractmethod


class Equation(ABC):
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
--fcc-editable-region--
```
