---
id: 662fa2e2cf27c09f21f4f5d0
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

في Python، تُعرف أنواع البيانات أثناء وقت التشغيل (عند تنفيذ الشفرة). لذلك، لا تحتاج إلى تحديد نوع البيانات للمتغير عند إعلانه. مع ذلك، يمكنك توضيح نوع البيانات التي سيحتويها المتغير باستخدام `variable: <data type> = value` أو فقط `variable: <data type>`. لاحظ أن مفسر Python لا يفرض أنواع البيانات المستخدمة في توضيح المتغيرات، وعادةً تحتاج إلى أدوات خارجية للقيام بذلك.

داخل صنف `Equation`، عرّف خاصية صنف باسم `degree`. لا تُعيّن لها قيمة. بدلاً من ذلك، استخدم توضيح نوع `int` لتبيين أنها ستخزن عددًا صحيحًا داخل الأصناف الملموسة.

لاحقًا، ستستخدم هذه الخاصية كجزء من عملية التحقق من المعلمات التي تُمرر لإنشاء كائنات المعادلة.

# --hints--

يجب أن تعرّف خاصية صنف باسم `degree` وتوضحها باستخدام `int` داخل صنف `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_variable("degree").is_equivalent("degree: int")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    def __init__(self):
        pass
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass
--fcc-editable-region--
lin_eq = LinearEquation()
```
