---
id: 662fbcef5f05e1b84f541a0c
title: الخطوة 13
challengeType: 20
dashedName: step-13
---

# --description--

سيتم إنشاء كل كائن معادلة بتمرير عدد من المعلمات يساوي عدد معاملات المعادلة، بدءًا من الدرجة n لـ \\( x \\) نزولًا إلى الدرجة الصفرية، بما في ذلك المعامل المحتمل بقيمة `0`.

على سبيل المثال، `LinearEquation(4, 5)` سيمثل المعادلة \\( 4x + 5 = 0 \\)، حيث `4` هو معامل الدرجة الأولى (الأعلى هنا) و`5` هو معامل الدرجة الصفرية.

يجب عليك التحقق من تمرير العدد الصحيح من المعلمات لإنشاء كائن المعادلة.

داخل طريقة `__init__`، أنشئ جملة `if` للتحقق مما إذا كان طول `args` مختلفًا عن عدد المعاملات التي يجب أن تحتويها المعادلة (`degree + 1`). إذا كان كذلك، ارفع استثناء `TypeError` واستخدم النص التالي لتوفير رسالة مخصصة: `f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"`.

ثم أصلح الخطأ بتمرير `2` و`3` لإنشاء `lin_eq`.

# --hints--

يجب أن تنشئ جملة `if` تتحقق مما إذا كان عدد المعاملات المستخدمة لإنشاء المعادلة مختلفًا عن `degree + 1`.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_conditions()[0]
cond.is_equivalent("(self.degree + 1) != len(args)") or cond.is_equivalent("len(args) != (self.degree + 1)")
`)) })
```

يجب أن ترفع استثناء `TypeError` داخل جملة `if` الجديدة وتستخدم النص المقدم لإرجاع رسالة خطأ مخصصة.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError(f"\\'{self.__class__.__name__}\\' object takes {self.degree + 1} positional arguments but {len(args)} were given")')
`)) })
```

يجب أن تمرر `2` و`3` لإنشاء `lin_eq`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation(2, 3)")`)) })
```


# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        pass        
    
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation()
--fcc-editable-region--
```
