---
id: 662fc3eba556a6bf800d48c1
title: الخطوة 14
challengeType: 20
dashedName: step-14
---

# --description--

الدالة المدمجة `isinstance()` تأخذ معلمتين وتُرجع قيمة منطقية تشير إلى ما إذا كان الكائن الممرر كالمعلمة الأولى هو نسخة من الصنف الممرر كالمعلمة الثانية.

```py
isinstance(7, int) # True
```

شيء آخر تريد التحقق منه هو أن كل معلمة رقم. بعد أول جملة `if`، أنشئ حلقة `for` تقوم بتكرار `args` وتتحقق مما إذا كانت المعلمة في التكرار الحالي ليست نسخة من `int` أو `float`. استخدم الدالة `isinstance()` ومرر لها صف مرتب يحتوي على `int` و `float` كالمعلمة الثانية.

إذا لم تكن المعلمة رقمًا، ارفع استثناء `TypeError` مع الرسالة `"Coefficients must be of type 'int' or 'float'"`.

# --hints--

يجب أن تنشئ حلقة `for` تقوم بتكرار `args` بعد جملة `if` الخاصة بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_iter().is_equivalent("args")`)) })
```

يجب أن تنشئ جملة `if` تتحقق مما إذا كان المعامل الحالي ليس نسخة من `int` أو `float` داخل حلقة `for`.

```js
({ test: () => assert(runPython(`
var = str(_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_vars())
cond1 = f'not isinstance({var}, (int, float))'
cond2 = f'not isinstance({var}, (float, int))'
if_stmt = _Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_conditions()[0]
if_stmt.is_equivalent(cond1) or if_stmt.is_equivalent(cond2)
`)) })
```

يجب أن تستخدم السلسلة النصية المقدمة لرفع استثناء `TypeError` داخل جملة `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_bodies()[0].has_stmt("raise TypeError(\\"Coefficients must be of type 'int' or 'float'\\")")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
```
