---
id: 6639fdcc701833a54c364211
title: الخطوة 17
challengeType: 20
dashedName: step-17
---

# --description--

بعد التحقق من صحة المعاملات، تحتاج إلى تخزينها في خاصية داخل الكائن. استخدم تعبير فهم القاموس لإنشاء قاموس يكون فيه المفتاح هو درجة المعامل، والقيمة المقابلة هي المعامل، ثم عيّن هذا القاموس إلى خاصية باسم `coefficients`.

على سبيل المثال، يجب أن يحتوي كائن `LinearEquation` الذي تم إنشاؤه باستخدام `2` و`4` على الخاصية `coefficients` التالية: `{1: 2, 0: 4}`، لأن `2` تمثل الدرجة الأولى من `x` و`4` تمثل الدرجة الصفرية من `x`.

أنشئ أزواج المفتاح-القيمة في القاموس الجديد باتباع نفس الترتيب الموجود في `args`.

# --hints--

يجب إعلان خاصية باسم `coefficients` داخل دالة `__init__` الخاصة بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_variable("self.coefficients")`)) })
```

يجب استخدام تعبير فهم القاموس لتخزين المعاملات.

```js
({ test: () => runPython(`
import ast
node = _Node(_code).find_class("Equation").find_function("__init__").find_variable("self.coefficients")
assert isinstance(node.tree.value, ast.DictComp)
`) })
```

يجب أن تكون خاصية `coefficients` قاموسًا يحتوي على أزواج المفتاح-القيمة بصيغة الدرجة-المعامل. تذكر اتباع نفس الترتيب الذي تُخزن به المعاملات داخل `args`.

```js
({ test: () => runPython(`
actual1 = list(LinearEquation(1, 6).coefficients.items())
expected1 = list({1: 1, 0: 6}.items())
actual2 = list(LinearEquation(-3.5, 0).coefficients.items())
expected2 = list({1: -3.5, 0: 0}.items())
assert actual1 == expected1
assert actual2 == expected2
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
--fcc-editable-region--
        
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass

    
lin_eq = LinearEquation(2, 3)
```
