---
id: 663b95d65caeb3ca04c5fef4
title: خطوة 21
challengeType: 20
dashedName: step-21
---

# --description--

إذا كان للمعامل قيمة غير صفرية، يمكنك أن تواجه حالات مختلفة. إذا كان `n == 0`، فإن الحد يتكون من المعامل نفسه.

بعد جملة `if` الخاصة بك، أنشئ جملة `if` أخرى لهذه الحالة وأضف سلسلة نصية تحتوي المعامل إلى قائمة `terms`. استخدم سلسلة f-string لذلك.

# --hints--

يجب أن تنشئ جملة `if` للتحقق مما إذا كانت قيمة `n` تساوي `0` بعد جملة `if` الموجودة لديك.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[1].find_conditions()[0].is_equivalent("n==0")`)) })
```

يجب أن تضيف `f'{coefficient}'` إلى قائمة `terms` داخل جملة `if` الجديدة.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[1].find_bodies()[0].is_equivalent("terms.append(f'{coefficient}')")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
        equation_string = ' '.join(terms)
        return equation_string
--fcc-editable-region--    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
