---
id: 663b83a28943e6aa6275a514
title: الخطوة 19
challengeType: 20
dashedName: step-19
---

# --description--

لا تزال داخل `Equation` الصنف، عرّف طريقة `__str__` لتوفير تمثيل نصي مناسب لكائنات المعادلة التي ستنشئها.

في الوقت الحالي، داخل طريقة `__str__`، أعلن متغيرًا `terms` وعيّنه إلى قائمة فارغة. ستستخدم هذا المتغير لتخزين كل حد (المعامل مضروبًا في \\( x^n \\)) من معادلتك.

بعد ذلك، أعلن متغيرًا `equation_string`، وعيّنه إلى نتيجة دمج العناصر في قائمة `terms` مع فراغ. أخيرًا، أرجع `equation_string`.

# --hints--

يجب أن تعرّف طريقة `__str__` داخل صنف `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__str__")`)) })
```

يجب أن تأخذ طريقة `__str__` معلمة واحدة، وهي `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_args("self")`)) })
```

يجب أن تعلن متغيرًا `terms` وتعيّنه إلى قائمة فارغة داخل طريقة `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("terms = []")`)) })
```

يجب أن تعلن متغيرًا `equation_string` وتعيّنه إلى نتيجة دمج العناصر في `terms` مع فراغ داخل طريقة `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("equation_string = ' '.join(terms)")`)) })
```

يجب أن تُرجع `equation_string` من طريقة `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_return("equation_string")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    
--fcc-editable-region--    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
