---
id: 6675a38a8b535e4ff3274520
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

تأخذ الدالة المدمجة `hasattr` كائنًا كمعلمتها الأولى وسلسلة نصية تمثل اسم خاصية كمعلمتها الثانية. تُرجع قيمة منطقية تشير إلى ما إذا كان الكائن يحتوي على الخاصية المحددة.

الآن ستستخدم الطريقة `__init_subclass__` للتحقق مما إذا كان الصنف الفرعي يحتوي على الخاصية `degree` عند لحظة التهيئة.

أنشئ عبارة `if` للتحقق مما إذا لم يكن لدى `cls` الخاصية `degree`. إذا كان الأمر كذلك، ارفع استثناء `AttributeError` واستخدم السلسلة النصية `f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"` لتوفير رسالة مخصصة.

بعد ذلك، أصلح الخطأ الذي ظهر في الطرفية بإعلان خاصية صنف `degree` داخل صنف `LinearEquation`. يجب أن تمثل هذه الخاصية درجة المعادلة، وهي الأس لأعلى حد \( x \). لذلك، عيّن العدد الصحيح `1` للخاصية `degree`.

# --hints--

يجب أن تنشئ عبارة `if` تتحقق مما إذا لم يكن لدى `cls` الخاصية `degree` داخل الطريقة `__init_subclass__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").find_ifs()[0].find_conditions()[0].is_equivalent("not hasattr(cls, 'degree')")`)) })
```

يجب أن ترفع استثناء `AttributeError` باستخدام السلسلة النصية المقدمة داخل عبارة `if` الخاصة بك.

```js
({ test: () => runPython(`
raise_stmt = 'raise AttributeError(f"Cannot create \\'{cls.__name__}\\' class: missing required attribute \\'degree\\'")'
node = _Node(_code).find_class("Equation").find_function("__init_subclass__").find_ifs()[0].find_bodies()[0]
assert node.has_stmt(raise_stmt) 
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        pass

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    
--fcc-editable-region--
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
