---
id: 6675aaf418b41157f6ccd692
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

حان الوقت للعودة إلى طريقة `__init__`. اعتمادًا على نوع المعادلة، ستحتاج إلى تمرير عدد متغير من المعلمات أثناء التهيئة.

أضف معلمة ثانية `args` إلى الطريقة واستخدم العامل `*` لجعلها تقبل عددًا متغيرًا من المعلمات.

# --hints--

يجب أن تأخذ طريقة `__init__` معلمتين، `self` و`*args`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self, *args")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    degree = 1

    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
