---
id: 66759e32b88fb5459b1e0234
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

تُستدعى طريقة `__init_subclass__` في كل مرة يُشتق فيها صنف من الصنف الذي يُعرّفها، وهي تتيح تخصيص الأصناف الفرعية. تأخذ الطريقة معلمة باسم `cls` حسب الاصطلاح (تمثّل "الصنف"), والتي تمثل الصنف الفرعي الجديد.

عرّف طريقة `__init_subclass__` في صنف `Equation` وأعطها معلمة `cls`.

# --hints--

يجب أن تعرف طريقة `__init_subclass__` بمعلمة `cls` في صنف `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").has_args("cls")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    
--fcc-editable-region--
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
