---
id: 66b59b12a745e10011158f7b
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

The next step is to create a few more variables that will be used in future bot messages.

Create a variable called `subject` and assign it the string value `"JavaScript"`.

Then create a variable called `topic` and assign it the string value `"strings"`.

# --hints--

Ви повинні мати змінну під назвою `subject`.

```js
assert.isNotNull(subject);
```

Призначте рядок до змінної `subject`.

```js
assert.isString(subject);
```

Призначте рядок `"JavaScript"` до змінної `subject`.

```js
assert.equal(subject, "JavaScript");
```

Ви повинні мати змінну під назвою `topic`.

```js
assert.isNotNull(topic);
```

Призначте рядок до змінної `topic`.

```js
assert.isString(topic);
```

Призначте рядок `"strings"` до змінної `topic`.

```js
assert.equal(topic, "strings");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

--fcc-editable-region--

--fcc-editable-region--
```
