---
id: 69bab8292750d84cac5b52a0
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

在這個工作坊中，吞吐量指的是在單次模擬（早上或晚上）中，所有時刻內每個閘門處理的總參加者數量。你的 `initializeThroughput` 函式（程式）將準備一個物件，用來追蹤每個閘門處理了多少參加者。

首先，創建一個名為 `summary` 的空物件。

# --hints--

在你的 `initializeThroughput` 函式（程式）中，創建一個空物件並指定給名為 `summary` 的變數。

```js
const funcStr = __helpers.removeJSComments(initializeThroughput.toString());
assert.match(
  funcStr,
  /(var|let|const)\s+summary/,
  "You must have at least one space between the declaration keyword (let/const) and 'summary'"
);
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(initializeThroughput.toString()));
const viable_solutions = [
"functioninitializeThroughput(gates){letsummary={};}",
"functioninitializeThroughput(gates){letsummary={}}",
"functioninitializeThroughput(gates){constsummary={};}",
"functioninitializeThroughput(gates){constsummary={}}",
"functioninitializeThroughput(gates){varsummary={};}",
"functioninitializeThroughput(gates){varsummary={}}"
];
assert.include(viable_solutions, cleaned);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
--fcc-editable-region--

--fcc-editable-region--
}
```
