---
id: 69babfc35c3c58d0993d404d
title: 第 9 步驟
challengeType: 1
dashedName: step-9
---

# --description--

為了完成你的 `processGateFlow` 函式，請傳回一個包含兩個屬性的物件：

- `processed`：該時脈內閘門能處理的參加者數量。
- `overflow`：處理後佇列中剩餘的參加者數量（即迴圈結束後的 `currentTickQueue`）。

你應該直接傳回該物件，如下所示：

```js
return {
  firstKey: firstValue,
  secondKey: secondValue
}
```

# --hints--

你應該傳回一個具有 `processed` 和 `overflow` 屬性的物件。直接傳回該物件，不要先指定給變數。

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(processGateFlow.toString()));
const viableSolutions = [
  "return{processed:processed,overflow:currentTickQueue}}",
  "return{processed:processed,overflow:currentTickQueue,}}",
  "return{overflow:currentTickQueue,processed:processed,}}",
  "return{overflow:currentTickQueue,processed:processed}}",
  "return{processed:processed,overflow:currentTickQueue};}",
  "return{processed:processed,overflow:currentTickQueue,};}",
  "return{overflow:currentTickQueue,processed:processed,};}",
  "return{overflow:currentTickQueue,processed:processed};}",
];
const endsCorrectly = viableSolutions.some(sol => cleaned.endsWith(sol));
assert(
  endsCorrectly,
  "Your processGateFlow function must return the object with processed and overflow properties."
);

const northGate = morningGates[0];
const result = processGateFlow(northGate, 0);
assert(
  typeof result === 'object' && result !== null && !Array.isArray(result),
  'Your function should return an object'
);
assert.hasAllKeys(
  result,
  ['processed', 'overflow'],
  'Your returned object should have properties `processed` and `overflow`'
);
```

你傳回的物件中，`processed` 屬性的值應該是 `processed`。

```js
const northGate = morningGates[0];
const result = processGateFlow(northGate, 0);
assert.strictEqual(
  result.processed,
  3,
  'For the North gate at tick 0 with capacity 5 and queue 3, processed should be 3'
);
const westGate = morningGates[3];
const resultWest = processGateFlow(westGate, 0);
assert.strictEqual(
  resultWest.processed,
  2,
  'For the West gate at tick 0 with capacity 2 and queue 4, processed should be 2'
);
```

你傳回的物件中，`overflow` 屬性的值應該是 `currentTickQueue`。

```js
const northGate = morningGates[0];
const result = processGateFlow(northGate, 0);
assert.strictEqual(
  result.overflow,
  0,
  'For the North gate at tick 0 with capacity 5 and queue 3, overflow should be 0'
);
const westGate = morningGates[3];
const resultWest = processGateFlow(westGate, 0);
assert.strictEqual(
  resultWest.overflow,
  2,
  'For the West gate at tick 0 with capacity 2 and queue 4, overflow should be 2'
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
