---
id: 69babfd65c4c4632a0c76f63
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

當一個閘口無法在一個時脈內處理所有參加者時，部分參加者會留在佇列中。為了控制這種上限溢位(overflow)，你將建置一個函式，將溢位的參加者重新導向到另一個閘口。

創建一個空的函式，名稱為 `rerouteOverflow`，並帶有以下參數：

- `gates`：包含所有閘口物件的完整陣列。
- `currentGate`：發生溢位的閘口。
- `tickIndex`：佇列中目前的時脈位置。
- `overflowAmount`：無法被處理的參加者數量。

# --hints--

你應該創建一個名為 `rerouteOverflow` 的函式。

```js
assert.isFunction(rerouteOverflow);
```

你的 `rerouteOverflow` 函式應該帶有參數 `gates`、`currentGate`、`tickIndex` 和 `overflowAmount`。

```js
const regex = __helpers.functionRegex('rerouteOverflow', ['gates', 'currentGate', 'tickIndex', 'overflowAmount']);
assert.match(__helpers.removeJSComments(rerouteOverflow.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

--fcc-editable-region--

--fcc-editable-region--
```
