---
id: 69babdf3f896711c11516373
title: 第 4 步驟
challengeType: 1
dashedName: step-4
---

# --description--

在每個模擬刻度中：

- 一定數量的參加者會從其 `queue` 抵達閘門。
- 閘門根據其 `capacity` 處理參加者。
- 如果抵達的參加者超過閘門能處理的數量，部分參加者將會留下（上限溢位）。

你現在將建置一個函式，來控制代碼這個邏輯，針對單一閘門在單一刻度的情況。

創建一個空的函式，名稱為 `processGateFlow`，並接受兩個參數：

- `gate`：正在處理的閘門物件。
- `tickIndex`：`queue` 陣列中目前模擬刻度的位置。

# --hints--

你應該創建一個名為 `processGateFlow` 的函式。

```js
assert.isFunction(processGateFlow);
```

你的函式 `processGateFlow` 應該有參數 `gate` 和 `tickIndex`。

```js
const regex = __helpers.functionRegex('processGateFlow', ['gate', 'tickIndex']);
assert.match(__helpers.removeJSComments(processGateFlow.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

--fcc-editable-region--

--fcc-editable-region--
```
