---
id: 69bac027df64d4433c708d44
title: 第 12 步驟
challengeType: 1
dashedName: step-12
---

# --description--

當重新導向上限溢位時，你應該將參加者送到 `gates` 陣列中的下一個 `gates`。為此，你需要找到下一個 `gate` 的索引。通常，你可以將目前的索引加 1。然而，如果目前的 `gate` 是陣列中的最後一個，你需要回繞到第一個 `gate`。

你可以使用取餘數運算子（`%`）來處理這個問題：

```js
const nextIndex = (currentIndex + 1) % array.length;
```

這樣做的原因是當 `currentIndex + 1` 等於 `array.length` 時，結果會變成 0。

使用這種方法，創建一個名為 `nextGateIndex` 的變數，用來存放 `gates` 陣列中下一個 `gate` 的索引。

# --hints--

你應該使用取餘數技巧，從 `currentIndex` 中存取並指定下一個索引到名為 `nextGateIndex` 的變數。請遵循範例中的同一行寫法。

```js
const rawFuncStr = __helpers.removeJSComments(rerouteOverflow.toString());
assert.match(
  rawFuncStr,
  /(var|const|let)\s+nextGateIndex/,
  "You must have at least one space between the declaration keyword (let/const) and 'nextGateIndex'."
);

const cleaned = __helpers.removeWhiteSpace(rawFuncStr);
const solution = /(var|const|let)nextGateIndex=\((currentIndex\+1|1\+currentIndex)\)%gates\.length;?}$/;

assert(
  solution.test(cleaned),
  "Inside your `rerouteOverflow` function, use the modulo technique to create a variable named `nextGateIndex` that stores the index of the next gate."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
--fcc-editable-region--

--fcc-editable-region--
}
```
