---
id: 69bac067cd15c4685d07906d
title: 第 15 步
challengeType: 1
dashedName: step-15
---

# --description--

現在你已經有函式可以處理單一刻度並控制上限溢位，接下來是為單一閘門創建一個控制元件函式。這個函式將在模擬的單一刻度中，處理該閘門的所有動作。

創建一個空的函式，名稱為 `handleGateAtTick`，並帶有以下參數：

- `gates`：完整的閘門物件陣列。
- `gate`：目前正在處理的閘門。
- `tickIndex`：模擬的當前刻度索引。
- `throughputSummary`：追蹤每個閘門總處理參加者數的物件。

# --hints--

你應該創建一個名為 `handleGateAtTick` 的函式。

```js
assert.isFunction(handleGateAtTick);
```

你的 `handleGateAtTick` 函式應該有參數 `gates`、`gate`、`tickIndex` 和 `throughputSummary`。

```js
const regex = __helpers.functionRegex('handleGateAtTick', ['gates', 'gate', 'tickIndex', 'throughputSummary']);
assert.match(__helpers.removeJSComments(handleGateAtTick.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

--fcc-editable-region--

--fcc-editable-region--
```
