---
id: 69baafd6d683dd9921b75db3
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將使用 JavaScript 模擬音樂節參加者的流動。

你會使用兩組描述音樂節大門在早班和晚班的資料集：

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];
```

每個大門物件包含以下屬性：

- `id`：一個用來識別大門的字串。
- `capacity`：大門每個刻度能處理的參加者數量。
- `queue`：一個數字陣列，代表在特定刻度期間抵達大門的參加者數量。

首先，創建一個名為 `initializeThroughput` 的空函式（程式），並帶有一個參數 `gates`。

# --hints--

你應該創建一個名為 `initializeThroughput` 的函式。

```js
assert.isFunction(initializeThroughput);
```

你的 `initializeThroughput` 函式（程式）應該帶有一個參數 `gates`。

```js
const explorer = await __helpers.Explorer(code);
const { initializeThroughput } = explorer.allFunctions;
const parameters = initializeThroughput?.parameters;
assert.lengthOf(parameters, 1);
assert.equal(parameters[0].toString(), "gates");
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

--fcc-editable-region--

--fcc-editable-region--
```
