---
id: 69d048937c7bb4b4c22ebb18
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

現在是時候讓參加者通過閘門了。在 `while` 迴圈內：

- 將 `currentTickQueue` 減少 1，表示有一位參加者已通過。使用遞減運算子（`--`）。
- 將 `processed` 增加 1，以追蹤處理過的參加者數量。使用累加運算子（`++`）。

# --hints--

你的 `while` 迴圈應該使用遞減運算子 `--` 對 `currentTickQueue` 進行遞減，並使用累加運算子 `++` 對 `processed` 進行累加。

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const viableSolutions = [
  "while(currentTickQueue>0&&processed<gate.capacity){processed++;currentTickQueue--;}}",
  "while(currentTickQueue>0&&processed<gate.capacity){processed++currentTickQueue--;}}",
  "while(currentTickQueue>0&&processed<gate.capacity){processed++;currentTickQueue--}}",
  "while(currentTickQueue>0&&processed<gate.capacity){processed++currentTickQueue--}}",
  "while(currentTickQueue>0&&processed<gate.capacity){currentTickQueue--;processed++;}}",
  "while(currentTickQueue>0&&processed<gate.capacity){currentTickQueue--processed++;}}",
  "while(currentTickQueue>0&&processed<gate.capacity){currentTickQueue--;processed++}}",
  "while(currentTickQueue>0&&processed<gate.capacity){currentTickQueue--processed++}}"
];
const endsCorrectly = viableSolutions.some(sol => cleaned.endsWith(sol));
assert(
  endsCorrectly,
  "Your `while` loop should contain both the `processed++` increment and the `currentTickQueue--` decrement, and nothing else."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
