---
id: 68e7e729f6d14f69a23786ef
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

بعد ذلك، داخل عنصر `main`، أضف عنصر `h1` يحتوي على النص `Parent Teacher Conference Form` والصنفين `title` و `center`.

# --hints--

يجب أن يكون لديك عنصر `h1`.

```js
assert.exists(document.querySelector("h1"));
```

يجب أن يحتوي عنصر `h1` على النص `Parent Teacher Conference Form`.

```js
assert.strictEqual(document.querySelector("h1")?.innerText, "Parent Teacher Conference Form");
```

يجب أن يحتوي عنصر `h1` على الصنفين `title` و `center`.

```js
assert.exists(document.querySelector("h1.title.center"));
```

يجب أن يكون عنصر `h1` داخل عنصر `main`.

```js
assert.exists(document.querySelector("main h1.title.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
--fcc-editable-region--
    
--fcc-editable-region--
  </main>
</body>

</html>
```
