---
id: 68e707257ea91c2e12d489fb
title: خطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستتدرب على كيفية إضافة أنماط مخصصة لأزرار الاختيار من خلال بناء نموذج مؤتمر أولياء الأمور والمعلمين. تم توفير قالب جاهز لـ HTML لك.

ابدأ بإضافة عنصر `main` مع صنف يسمى `container`.

# --hints--

يجب أن يكون لديك عنصر `main`.

```js
assert.exists(document.querySelector("main"));
```

يجب أن يحتوي عنصر `main` على صنف يسمى `container`.

```js
assert.exists(document.querySelector("main.container"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
