---
id: 68ea8a5db73b155708992687
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

الخطوة التالية هي إضافة عنصر `input` المرتبط بمعلومات الطالب.

ابدأ بإضافة عنصر `input` مع خاصية `type` مضبوطة على `"text"`. ثم أضف خاصيتي `name` و `id`، كلاهما مضبوطان على `student-name`.

بعد ذلك، أضف خاصية `placeholder` مضبوطة على `E.g., Jane Doe`. وأخيرًا، أضف خاصية `required`.

# --hints--

يجب أن يكون لديك عنصر `input`.

```js
assert.exists(document.querySelector("fieldset input"));
```

يجب أن يحتوي عنصر `input` على خاصية `type` مضبوطة على `"text"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

يجب أن يحتوي عنصر `input` على خاصية `name` مضبوطة على `"student-name"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("name"), "student-name");
```

يجب أن يحتوي عنصر `input` على خاصية `id` مضبوطة على `"student-name"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("id"), "student-name");
```

يجب أن يحتوي عنصر `input` على خاصية `placeholder` مضبوطة على `E.g., Jane Doe`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Jane Doe");
```

يجب أن يحتوي عنصر `input` على خاصية `required`.

```js
const inputEl = document.querySelector("fieldset input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>

      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>


    </form>
  </main>
</body>

</html>
```
