---
id: 68ea8fdf2c9b167307e7f671
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

حان الوقت الآن لإضافة عناصر النموذج لجمع معلومات الوالدين.

ابدأ بإضافة عنصر `label` مع النص `Parent/Guardian Name: ` وخاصية `for` مضبوطة على `"parent-name"`.

ثم، أسفل عنصر `label` الخاص بك، أضف عنصر `input` مع خاصية `type` مضبوطة على `"text"`. يجب ضبط خاصيتي `name` و `id` على `"parent-name"`. يجب ضبط خاصية `placeholder` على `"E.g., Nancy Doe"`. وأخيرًا، يجب أن يكون عنصر `input` الخاص بك مطلوبًا.

# --hints--

يجب أن يكون لديك ثلاثة عناصر `label` في الصفحة.

```js
assert.lengthOf(document.querySelectorAll("label"), 3);
```

يجب أن يحتوي عنصر `label` الثالث على النص `Parent/Guardian Name: `. لا تنسَ المسافة بعد النقطتين.

```js
assert.match(code, /Parent\/Guardian\s+Name:\s+<\/label>/);
```

يجب أن يحتوي عنصر `label` الثالث على خاصية `for`.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

يجب أن تكون قيمة خاصية `for` في عنصر `label` الثالث هي `"parent-name"`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type label")?.getAttribute("for"), "parent-name");
```

يجب أن يكون لديك ثلاثة عناصر `input` في الصفحة.

```js
assert.lengthOf(document.querySelectorAll("input"), 3);
```

يجب أن يكون عنصر `input` الثالث أسفل عنصر `label` الثالث.

```js
const fieldsetChildren = [...document.querySelector("fieldset:last-of-type")?.children];
const inputEl = document.querySelector("fieldset:last-of-type input");
const labelEl = document.querySelector("fieldset:last-of-type label");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

يجب أن يحتوي عنصر `input` الثالث على خاصية `type` مضبوطة على `"text"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

يجب أن يحتوي عنصر `input` الثالث على خاصية `name` مضبوطة على `"parent-name"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("name"), "parent-name");
```

يجب أن يحتوي عنصر `input` الثالث على خاصية `id` مضبوطة على `"parent-name"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("id"), "parent-name");
```

يجب أن يحتوي عنصر `input` الثالث على خاصية `placeholder` مضبوطة على `"E.g., Nancy Doe"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Nancy Doe");
```

يجب أن يحتوي عنصر `input` الثالث على خاصية `required`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
