---
id: 68ea975c7cbbb894a386e402
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

الخطوة التالية هي إضافة عناصر `label` و `input` لطريقة الاتصال عبر البريد الإلكتروني.

ابدأ بإضافة مسمى (label) بصنف `"contact-method"` وخاصية `for` مضبوطة على `"email"`. يجب أن يكون نص المسمى `Email: `.

أسفل عنصر `label` الخاص بك، أضف زر اختيار (radio button) بخواص `id` و `value` مضبوطة على `"email"`. يجب ضبط خاصية `name` على `"contact-method"` والصنف على `"contact-method-radio-btn"`.

أخيرًا، تأكد من أن زر الاختيار هذا محدد (checked) افتراضيًا.

# --hints--


يجب أن تحتوي مجموعة الحقول (fieldset) الأخيرة على عنصر `label`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

يجب أن يحتوي عنصر `label` الجديد (للبريد الإلكتروني) على صنف `"contact-method"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

يجب ضبط خاصية `for` في عنصر `label` الجديد على `"email"`.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "email");
```

يجب أن يحتوي عنصر `label` على النص `Email: `. تأكد من وجود مسافة بعد النقطتين.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.match(label?.innerText, /Email:\s?/);
```

يجب أن يحتوي آخر `fieldset` على عنصر `input`.

```js
assert.exists(document.querySelector("fieldset:last-of-type input"));
```

يجب أن يحتوي عنصر `input` على خاصية `type` مضبوطة على `radio`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

يجب أن يحتوي عنصر `input` على خاصية `id` مضبوطة على `"email"`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("id"), "email");
```

يجب أن يحتوي عنصر `input` على خاصية `value` مضبوطة على `"email"`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("value"), "email");
```

يجب أن يحتوي عنصر `input` على خاصية `name` مضبوطة على `"contact-method"`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

يجب أن يحتوي عنصر `input` على صنف `"contact-method-radio-btn"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type input.contact-method-radio-btn"));
```

يجب أن يكون زر الاختيار (`radio`) محددًا (`checked`) افتراضيًا.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.isTrue(input?.hasAttribute("checked"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
