---
id: 68eaa668d4d9fbe529a0e16b
title: الخطوة 14
challengeType: 0
dashedName: step-14
---

# --description--

أضف الآن عنصر `label` بالنص `Any specific concerns or topics you'd like to discuss?` وخاصية `for` مضبوطة على `notes`.

أسفل هذا الـ `label`، أضف عنصر `textarea`. عيّن الخاصية `id` إلى `notes`، و`name` إلى `notes`، و`rows` إلى `4`، و`cols` إلى `50`.

# --hints--

يجب أن تحتوي آخر مجموعة حقول (`fieldset`) على عنصر `label`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

يجب أن يحتوي عنصر `label` على النص `Any specific concerns or topics you'd like to discuss?`.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.innerText, "Any specific concerns or topics you'd like to discuss?");
```

يجب ضبط خاصية `for` في عنصر `label` على `"notes"`.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "notes");
```

يجب أن تحتوي مجموعة الحقول (`fieldset`) الأخيرة على عنصر `textarea`.

```js
assert.exists(document.querySelector("fieldset:last-of-type textarea"));
```

يجب أن يكون عنصر `textarea` أسفل عنصر `label`.

```js
const label = document.querySelector("fieldset:last-of-type label");
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(label?.nextElementSibling, textarea);
```

يجب أن يحتوي عنصر `textarea` على خاصية `id` مضبوطة على `"notes"`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("id"), "notes");
```

يجب أن يحتوي عنصر `textarea` على خاصية `name` مضبوطة على `"notes"`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("name"), "notes");
```

يجب أن يحتوي عنصر `textarea` على خاصية `rows` بقيمة `4`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("rows"), "4");
```

يجب أن يحتوي عنصر `textarea` على خاصية `cols` بقيمة `50`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("cols"), "50");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
