---
id: 68eaa310c153c0d0c5c9b78e
title: الخطوة 13
challengeType: 0
dashedName: step-13
---

# --description--

من الممارسات الجيدة السماح للمستخدمين بإضافة ملاحظات أو مخاوف إضافية.

أسفل عنصر `fieldset` الثالث، أضف عنصر `fieldset` جديد. بداخله، أضف عنصر `legend` بالنص `Additional Notes`.


# --hints--

يجب أن يكون لديك إجمالي 4 عناصر `fieldset` في هذه الصفحة.

```js
const fieldsets = document.querySelectorAll('fieldset');
assert.lengthOf(fieldsets, 4);
```

يجب أن تضيف عنصر `fieldset` أسفل عنصر `fieldset` الثالث.

```js
const fieldsets = document.querySelectorAll('fieldset');
assert.exists(fieldsets[2]?.nextElementSibling?.matches('fieldset'));
```

يجب أن يحتوي عنصر `fieldset` الجديد على عنصر `legend`.

```js
const newFieldset = document.querySelectorAll('fieldset')[3];
assert.exists(newFieldset?.querySelector('legend'));
```

يجب أن يحتوي عنصر `legend` على النص `Additional Notes`.

```js
const newFieldset = document.querySelectorAll('fieldset')[3];
const legend = newFieldset?.querySelector('legend');
assert.equal(legend?.innerText?.trim(), 'Additional Notes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
