---
id: 68ea8666b8b473455f5b33ef
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

سيتركز القسم الأول من النموذج على معلومات الطالب، مثل اسمه ودرجته.

داخل عنصر `form`، أضف عنصر `fieldset`. داخل ذلك العنصر `fieldset`، أضف عنصر `legend` بالنص `Student Information`.

أسفل عنصر `legend`، أضف عنصر `label` بالنص `Full Name: ` وخاصية `for` بقيمة `student-name`.


# --hints--

يجب أن يكون لديك عنصر `fieldset`.

```js
assert.exists(document.querySelector("fieldset"));
```

يجب أن يكون عنصر `fieldset` داخل عنصر `form`.

```js
assert.exists(document.querySelector("form fieldset"));
```

يجب أن يكون لديك عنصر `legend`.

```js
assert.exists(document.querySelector("legend"));
```

يجب أن يكون عنصر `legend` داخل عنصر `fieldset`.

```js
assert.exists(document.querySelector("fieldset legend"));
```

يجب أن يحتوي عنصر `legend` على النص `Student Information`.

```js
assert.strictEqual(document.querySelector("fieldset legend")?.innerText, "Student Information");
```

يجب أن يكون لديك عنصر `label`.

```js
assert.exists(document.querySelector("label"));
```

يجب أن يحتوي عنصر `label` على النص `Full Name: `. لا تنسَ المسافة بعد النقطتين.

```js
assert.match(code, /Full\s+Name:\s+<\/label>/);
```

يجب أن يحتوي عنصر `label` على خاصية `for`.

```js
assert.isTrue(document.querySelector("label")?.hasAttribute("for"));
```

يجب أن تكون قيمة خاصية `for` في عنصر `label` هي `student-name`.

```js
assert.strictEqual(document.querySelector("label")?.getAttribute("for"), "student-name");
```

يجب أن يكون عنصر `label` أسفل عنصر `legend`.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const legendEl = document.querySelector("fieldset legend");
const labelEl = document.querySelector("label");

assert.isBelow(fieldsetChildren.indexOf(legendEl), fieldsetChildren.indexOf(labelEl));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
