---
id: 68ea8c330ed07b60879f3d38
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

حان الوقت الآن لإضافة عناصر النموذج لجمع معلومات درجة الطالب.

ابدأ بإضافة عنصر `label` آخر يحتوي على النص `Student Grade: ` وخاصية `for` مضبوطة على `"grade"`.

ثم، أسفل عنصر `label` الخاص بك، أضف عنصر `input` مع خاصية `type` مضبوطة على `"number"`. يجب ضبط خاصيتي `name` و `id` على `"grade"`. يجب ضبط خاصية `placeholder` على `"E.g., 4"`. وأخيرًا، يجب أن يكون عنصر `input` الخاص بالأرقام مطلوبًا.

# --hints--

يجب أن يكون لديك مجموع عنصرين `label` في الصفحة.

```js
assert.lengthOf(document.querySelectorAll("label"), 2);
```

يجب أن يحتوي عنصر `label` الثاني على النص `Student Grade: `. لا تنسَ المسافة بعد النقطتين.

```js
assert.match(code, /Student\s+Grade:\s+<\/label>/);
```

يجب أن يحتوي عنصر `label` الثاني على خاصية `for`.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

يجب أن تكون قيمة خاصية `for` في عنصر `label` الثاني هي `"grade"`.

```js
assert.strictEqual(document.querySelector("label:last-of-type")?.getAttribute("for"), "grade");
```

يجب أن يكون لديك عنصرين `input` في الصفحة.

```js
assert.lengthOf(document.querySelectorAll("input"), 2);
```

يجب أن يكون عنصر `input` الثاني أسفل عنصر `label` الثاني.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const inputEl = document.querySelector("fieldset input:last-of-type");
const labelEl = document.querySelector("label:last-of-type");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

يجب أن يحتوي عنصر `input` الثاني على خاصية `type` مضبوطة على `"number"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("type"), "number");
```

يجب أن يحتوي عنصر `input` الثاني على خاصية `name` مضبوطة على `"grade"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("name"), "grade");
```

يجب أن يحتوي عنصر `input` الثاني على خاصية `id` مضبوطة على `"grade"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("id"), "grade");
```

يجب أن يحتوي عنصر `input` الثاني على خاصية `placeholder` مضبوطة على `"E.g., 4"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., 4");
```

يجب أن يحتوي عنصر `input` الثاني على خاصية `required`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
