---
id: 68ea81da9f687732368a5322
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

حان الوقت الآن لإضافة عناصر `form` و `input`، التي ستمثل معلومات الوالد والطالب.

ابدأ بإضافة عنصر `form` أسفل عنصر `p`.

# --hints--

يجب أن يكون لديك عنصر `form`.

```js
assert.exists(document.querySelector(".container form"));
```

يجب أن يكون عنصر form أسفل عنصر `p`.

```js
assert.exists(document.querySelector(".container > p+form"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </main>
</body>

</html>
```
