---
id: 5895f70bf9fc0f352b528e64
title: 使用模板引擎的強大功能
challengeType: 2
forumTopicId: 301567
dashedName: use-a-template-engines-powers
---

# --description--

使用模板引擎的最大特性之一是能夠在將模板檔案呈現為 HTML 之前，從伺服器傳遞變數到模板檔案。

在你的 Pug 檔案中，你可以透過在元素的其他文字中行內以 `#{variable_name}` 取用變數名稱，或是在元素上使用不帶空格的等號，例如 `p=variable_name`，將變數的值指定給 p 元素的文字。

Pug 完全是透過使用空白和製表符來顯示巢狀元素，並減少製作美觀網站所需的程式碼量。

以以下 Pug 程式碼為例： 
 
```pug
head
  script(type='text/javascript').
    if (foo) bar(1 + 5);
body
  if youAreUsingPug
      p You are amazing
    else
      p Get on it!
```

上述產生以下 HTML：

```html
<head>
  <script type="text/javascript">
    if (foo) bar(1 + 5);
  </script>
</head>
<body>
  <p>You are amazing</p>
</body>
```

你專案中包含的 `index.pug` 檔案使用了變數 `title` 和 `message`。

將那些從你的伺服端傳遞到 Pug 檔案，方法是為你的 `res.render` 呼叫添加一個物件作為第二個引數，該物件包含變數及其值。給 `title` 賦值為 `Hello`，並給 `message` 賦值為 `Please log in`。

它應該看起來像這樣：

```javascript
res.render('index', { title: 'Hello', message: 'Please log in' });
```

現在重新整理你的頁面，你應該會在你的檢視中看到那些值，並且在你的 `index.pug` 檔案中所設定的正確位置呈現。

當你認為已經完成正確時，提交你的頁面。如果你遇到錯誤，可以查看<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#use-a-template-engines-power-2" target="_blank" rel="noopener noreferrer nofollow">至此為止完成的專案</a>。

# --hints--

Pug 應該正確地呈現變數。

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /pug-variable("|')>Please log in/gi,
    'Your projects home page should now be rendered by pug with the projects .pug file unaltered'
  );
```

