---
id: 5895f70cf9fc0f352b528e67
title: 實作 Passport 使用者的序列化
challengeType: 2
forumTopicId: 301556
dashedName: implement-the-serialization-of-a-passport-user
---

# --description--

你沒有載入實際的使用者物件，因為資料庫尚未設定。當你啟動伺服端時，連接一次資料庫，並在應用程式的整個生命週期中保持持久連線。為此，將你的資料庫連接字串（例如：`mongodb+srv://<username>:<password>@cluster0-jvwxi.mongodb.net/?retryWrites=true&w=majority`）添加到環境變數 `MONGO_URI`。這會在 `connection.js` 檔案中使用。

*如果你在設定 MongoDB Atlas 上的免費資料庫時遇到問題，請參考這個 <a href="https://www.freecodecamp.org/news/get-started-with-mongodb-atlas/" target="_blank" rel="noopener noreferrer nofollow">教學</a>。*

現在你想連接到你的資料庫，然後開始監聽請求。這樣做的目的是在你的資料庫連接之前或發生資料庫錯誤時，不允許請求。為了達成此目的，請將你的序列化和應用程式路由包裹在以下程式碼中：

```javascript
myDB(async client => {
  const myDataBase = await client.db('database').collection('users');

  // Be sure to change the title
  app.route('/').get((req, res) => {
    // Change the response to render the Pug template
    res.render('index', {
      title: 'Connected to Database',
      message: 'Please login'
    });
  });

  // Serialization and deserialization here...

  // Be sure to add this...
}).catch(e => {
  app.route('/').get((req, res) => {
    res.render('index', { title: e, message: 'Unable to connect to database' });
  });
});
// app.listen out here...
```

請務必取消註解 `deserializeUser` 中的 `myDataBase` 程式碼，並編輯你的 `done(null, null)` 以包含 `doc`。

當你認為已經正確完成時，提交你的頁面。如果你遇到錯誤，你可以<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implement-the-serialization-of-a-passport-user-5" target="_blank" rel="noopener noreferrer nofollow">取用至此為止完成的專案</a>。

# --hints--

資料庫連接應該存在。

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Connected to Database/gi,
    'You successfully connected to the database!'
  );
```

反序列化現在應該正確使用資料庫，並且應該以 `doc` 呼叫 `done(null, null)`。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /null,\s*doc/gi,
    'The callback in deserializeUser of (null, null) should be altered to (null, doc)'
  );
```

