---
id: 5895f700f9fc0f352b528e63
title: 設定模板引擎
challengeType: 2
forumTopicId: 301564
dashedName: set-up-a-template-engine
---

# --description--

處理這些挑戰將涉及你使用以下其中一種方法來撰寫你的程式碼：

- 複製<a href="https://github.com/freeCodeCamp/boilerplate-advancednode/" target="_blank" rel="noopener noreferrer nofollow">這個 GitHub 倉庫</a>並在區域端完成這些挑戰。
- 使用你選擇的網站建置工具來完成專案。務必將我們 GitHub 倉庫中的所有檔案納入其中。

模板引擎啟用你在應用程式中使用靜態模板檔案（例如以 *Pug* 撰寫的檔案）。在執行期，模板引擎會將模板檔案中的變數替換為由你的伺服器提供的實際值。然後，它會將模板轉換成靜態 HTML 檔案並傳送給客戶端。這種方法讓設計 HTML 頁面變得更簡單，並且允許在頁面上顯示變數，而不需要從客戶端進行 API 呼叫。

`pug@~3.0.0` 已經安裝，並且已列為你 `package.json` 檔案中的相依性。

Express 需要知道你正在使用哪個模板引擎。使用 `set` 方法將 `pug` 指定為 `view engine` 屬性的值：

```javascript
app.set('view engine', 'pug');
```

接著，為你的 `app` 添加另一個 `set` 方法，將 `views` 屬性設定為指向 `./views/pug` 目錄。這會告訴 Express 對所有視圖以該目錄為相對路徑來呈現。

最後，在你的首頁路由中使用 `res.render()`，並將 `index` 作為第一個引數傳入。這將呈現 `pug` 範本。

如果一切按計劃進行，你的應用程式首頁將不再是空白。取而代之的是，它會顯示一則訊息，指出你已成功呈現 `Pug template`！

當你認為已經完成時，提交你的頁面。如果你遇到錯誤，你可以<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-a-template-engine-1" target="_blank" rel="noopener noreferrer nofollow">查看至此為止完成的專案</a>。

# --hints--

Pug 應該是一個相依性。

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'pug',
    'Your project should list "pug" as a dependency'
  );
```

檢視引擎應該是 Pug。

```js
  const url = new URL("/_api/app", code);
  const res = await fetch(url);
  const app = await res.json();
  assert.equal(app?.settings?.['view engine'], "pug");
```

你應該將應用程式的 `views` 屬性設定為 `./views/pug`。

```js
  const url = new URL("/_api/app", code);
  const res = await fetch(url);
  const app = await res.json();
  assert.equal(app?.settings?.views, "./views/pug");
```

使用正確的 ExpressJS 方法從回應中呈現索引頁面。

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
      assert.match(
        data,
        /FCC Advanced Node and Express/gi,
        'You successfully rendered the Pug template!'
      );
```

Pug 應該可以正常運作。

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
      assert.match(
        data,
        /pug-success-message/gi,
        'Your projects home page should now be rendered by pug with the projects .pug file unaltered'
      );
```

