---
id: 5895f70df9fc0f352b528e68
title: 驗證策略
challengeType: 2
forumTopicId: 301547
dashedName: authentication-strategies
---

# --description--

策略是一種驗證使用者的方式。你可以使用策略來允許使用者根據區域儲存的資訊進行驗證（如果你先讓他們註冊）或來自各種提供者，例如 Google 或 GitHub。對於這個專案，我們將使用 Passport 中介層。Passport 提供一套完整的策略，支援使用使用者名稱和密碼、GitHub、Google 等方式進行驗證。

`passport-local@~1.0.0` 已經被添加為相依性。將它添加到你的伺服器如下：

```javascript
const LocalStrategy = require('passport-local');
```

告訴 passport **使用** 一個具現化的 `LocalStrategy` 物件並定義一些設定。請確保這個（以及從此點開始的所有內容）都封裝在資料庫連接中，因為它依賴於該連接！

```javascript
passport.use(new LocalStrategy((username, password, done) => {
  myDataBase.findOne({ username: username }, (err, user) => {
    console.log(`User ${username} attempted to log in.`);
    if (err) return done(err);
    if (!user) return done(null, false);
    if (password !== user.password) return done(null, false);
    return done(null, user);
  });
}));
```

這是在定義當你嘗試區域驗證某人時要使用的處理程序。首先，它會嘗試在你的資料庫中尋找輸入的使用者名稱。接著，它會檢查密碼是否匹配。最後，如果沒有出現你檢查過的錯誤（例如密碼錯誤），則會傳回 `user` 物件，並且該使用者已通過驗證。

許多策略是使用不同設定來建立的。通常，根據該策略的儲存庫中的 README 設定它很簡單。一個很好的範例是 GitHub 策略，你不需要擔心使用者名稱或密碼，因為使用者會被導向 GitHub 的認證頁面進行驗證。只要他們已登入並同意，GitHub 就會傳回他們的評測供你使用。

在下一步，你將設定如何實際呼叫驗證策略，以根據表單資料驗證使用者。

當你認為已經正確完成時，提交你的頁面。如果你遇到錯誤，可以<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#authentication-strategies-6" target="_blank" rel="noopener noreferrer nofollow">查看至此為止完成的專案</a>。

# --hints--

Passport-local 應該是一個相依性。

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport-local',
    'Your project should list "passport-local " as a dependency'
  );
```

應正確引入並設定 Passport-local。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')passport-local("|')/,
    'You should have required passport-local'
  );
  assert.match(
    data,
    /new LocalStrategy/,
    'You should have told passport to use a new strategy'
  );
  assert.match(
    data,
    /findOne/,
    'Your new local strategy should use the findOne query to find a username based on the inputs'
  );
```

