---
id: 5895f70df9fc0f352b528e6a
title: 創建新的中介層
challengeType: 2
forumTopicId: 301551
dashedName: create-new-middleware
---

# --description--

如同現在，任何使用者只要輸入 URL 就可以直接前往 `/profile`，不論他們是否已通過驗證。你想要透過先檢查使用者是否已驗證，然後再呈現評測頁面來防止這種情況。這正是創建中介層的完美範例。

這裡的挑戰是建立中介層函式 `ensureAuthenticated(req, res, next)`，它會透過在 `request` 上呼叫 Passport 的 `isAuthenticated` 方法來檢查使用者是否已驗證，該方法會檢查 `req.user` 是否已定義。如果是，則應該呼叫 `next()`。否則，你可以直接回應該請求並重新導向到你的首頁以進行登入。

這個中介層的實作為：

```javascript
function ensureAuthenticated(req, res, next) {
  if (req.isAuthenticated()) {
    return next();
  }
  res.redirect('/');
};
```

創建上述中介層函式，然後在 GET 請求的引數之前，將 `ensureAuthenticated` 作為中介層傳遞給評測頁面的請求：

```javascript
app
 .route('/profile')
 .get(ensureAuthenticated, (req,res) => {
    res.render('profile');
 });
```

當你認為頁面已經正確時，請提交你的頁面。如果你遇到錯誤，你可以<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#create-new-middleware-8" target="_blank" rel="noopener noreferrer nofollow">查看至此為止完成的專案</a>。

# --hints--

應該實作中介層 `ensureAuthenticated` 並將其附加到 `/profile` 路由。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /ensureAuthenticated[^]*req.isAuthenticated/,
    'Your ensureAuthenticated middleware should be defined and utilize the req.isAuthenticated function'
  );
  assert.match(
    data,
    /profile[^]*get[^]*ensureAuthenticated/,
    'Your ensureAuthenticated middleware should be attached to the /profile route'
  );
```

未經驗證的 GET 請求到 `/profile` 應正確重新導向到 `/`。

```js
  const url = new URL("/profile", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Home page/,
    'An attempt to go to the profile at this point should redirect to the homepage since we are not logged in'
  );
```

