---
id: 5895f70cf9fc0f352b528e66
title: 使用者物件的序列化
challengeType: 2
forumTopicId: 301563
dashedName: serialization-of-a-user-object
---

# --description--

序列化和反序列化是在認證方面的重要概念。將物件序列化表示將其內容轉換成一個小的 *key*，然後可以將該 *key* 反序列化回原始物件。這使我們能夠知道誰與伺服器通訊，而不必在每次請求新頁面時傳送認證資料，例如使用者名稱和密碼。

要正確設定這個，你需要有一個 serialize 函式和一個 deserialize 函式。在 Passport 中，這些可以用以下方式創建：

```javascript
passport.serializeUser(cb);
passport.deserializeUser(cb);
```

傳遞給 `serializeUser` 的回呼函式會以兩個引數被呼叫：完整的使用者物件，以及 passport 使用的回呼。 

回呼預期兩個引數：如果有錯誤，則為錯誤，以及用來識別應該在回呼中傳回的使用者的唯一鍵。你將在物件中使用使用者的 `_id`。這是保證唯一的，因為它是由 MongoDB 產生的。

同樣地，`deserializeUser` 以兩個引數呼叫：唯一鍵和回呼函式。

此回呼預期兩個引數：若有錯誤，則為錯誤，以及完整的使用者物件。若要取得完整的使用者物件，請對 Mongo 的 `_id` 進行查詢搜尋，如下所示：


```javascript
passport.serializeUser((user, done) => {
  done(null, user._id);
});

passport.deserializeUser((id, done) => {
  myDataBase.findOne({ _id: new ObjectID(id) }, (err, doc) => {
    done(null, null);
  });
});
```

將上述兩個函式（程式）添加到你的伺服端。`ObjectID` 類別來自 `mongodb` 軟體包。`mongodb@~3.6.0` 已經被添加為相依項。使用以下方式宣告此類別：

```javascript
const { ObjectID } = require('mongodb');
```

`deserializeUser` 會發出錯誤，直到你設定好資料庫連線。所以，暫時將 `myDatabase.findOne` 呼叫註解掉，並且只在 `deserializeUser` 回呼函式中呼叫 `done(null, null)`。

當你認為已經完成頁面時，請提交。如果你遇到錯誤，可以<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#serialization-of-a-user-object-4" target="_blank" rel="noopener noreferrer nofollow">查看至此為止完成的專案</a>。

# --hints--

你應該正確序列化使用者物件。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport.serializeUser/gi,
    'You should have created your passport.serializeUser function'
  );
  assert.match(
    data,
    /null,\s*user._id/gi,
    'There should be a callback in your serializeUser with (null, user._id)'
  );
```

你應該正確地反序列化使用者物件。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport.deserializeUser/gi,
    'You should have created your passport.deserializeUser function'
  );
  assert.match(
    data,
    /null,\s*null/gi,
    'There should be a callback in your deserializeUser with (null, null) for now'
  );
```

MongoDB 應該是一個相依性。

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'mongodb',
    'Your project should list "mongodb" as a dependency'
  );
```

Mongodb 應該正確地被 require，包括 `ObjectId`。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')mongodb\1/gi,
    'You should have required mongodb'
  );
  assert.match(
    data,
    /new ObjectID.*id/gi,
    'Even though the block is commented out, you should use new ObjectID(id) for when we add the database'
  );
```

