---
id: 5895f70cf9fc0f352b528e65
title: 設定 Passport
challengeType: 2
forumTopicId: 301565
dashedName: set-up-passport
---

# --description--

現在是設定 *Passport* 的時候了，這樣你就可以開始允許使用者註冊或登入帳號。除了 Passport，你還會使用 Express-session 來控制代碼會話。Express-session 有許多進階功能可以使用，但目前你只會使用基本功能。使用這個中介層會將會話 ID 以 cookie 形式存放在客戶端，並允許我們在伺服端使用該 ID 存取會話資料。這樣一來，你就能將個人帳號資訊從客戶端用來告訴伺服端客戶端已通過驗證的 cookie 中分離，並保留存取伺服端資料的 *key*。

`passport@~0.4.1` 和 `express-session@~1.17.1` 已經安裝，並且都列在你的 `package.json` 檔案中的依賴項。

你需要設定會話設定並初始化 Passport。首先，創建變數 `session` 和 `passport`，分別用來 require `express-session` 和 `passport`。

然後，透過定義以下選項，設定你的 Express 應用程式以使用 session：

```javascript
app.use(session({
  secret: process.env.SESSION_SECRET,
  resave: true,
  saveUninitialized: true,
  cookie: { secure: false }
}));
```

務必將 `SESSION_SECRET` 添加到你的 `.env` 檔案，並給它一個隨機值。這用於計算用來加密你的 cookie 的雜湊值！

在你完成所有這些之後，告訴你的 express 應用程式 **use** `passport.initialize()` 和 `passport.session()`。

當你認為已經正確完成時，提交你的頁面。如果你遇到錯誤，你可以<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-passport-3" target="_blank" rel="noopener noreferrer nofollow">查看至此為止完成的專案</a>。

# --hints--

Passport 和 Express-session 應該是相依套件。

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport',
    'Your project should list "passport" as a dependency'
  );
  assert.property(
    packJson.dependencies,
    'express-session',
    'Your project should list "express-session" as a dependency'
  );
```

應正確 require 依賴項。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')passport("|')/,
    'You should have required passport'
  );
  assert.match(
    data,
    /require.*("|')express-session("|')/,
    'You should have required express-session'
  );
```

Express 應用程式應該使用新的相依性。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(data, /passport\.initialize/, 'Your express app should use "passport.initialize()"');
  assert.match(data, /passport\.session/, 'Your express app should use "passport.session()"');
```

應正確設定 Session 和 session secret。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /secret *:\s*process\.env(\.SESSION_SECRET|\[(?<q>"|')SESSION_SECRET\k<q>\])/,
    'Your express app should have express-session set up with your secret as process.env.SESSION_SECRET'
  );
```

