---
id: 5895f70ef9fc0f352b528e6b
title: 如何組合一個評測
challengeType: 2
forumTopicId: 301554
dashedName: how-to-put-a-profile-together
---

# --description--

既然你可以確保存取 `/profile` 的使用者已通過驗證，你就可以在你的頁面上使用 `req.user` 中包含的資訊。

將包含屬性 `username` 且值為 `req.user.username` 的物件作為第二個引數，傳遞給評測視圖的 `render` 方法。

然後，前往你的 `profile.pug` 檢視，並在現有的 `h1` 元素下方、相同縮排層次，添加以下這一行：

```pug
h2.center#welcome Welcome, #{username}!
```

這會創建一個帶有類別 `center` 和 id `welcome` 的 `h2` 元素，包含文字 `Welcome, `，後面接使用者名稱。

此外，在 `profile.pug` 中，為 `/logout` 路由添加一個取用該路由的聯結，該路由將承載取消使用者認證的邏輯：

```pug
a(href='/logout') Logout
```

當你認為已經完成時，提交你的頁面。如果你遇到錯誤，你可以<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#how-to-put-a-profile-together-9" target="_blank" rel="noopener noreferrer nofollow">查看至此為止完成的專案</a>。

# --hints--

你應該正確地將 Pug 呈現變數添加到 `/profile`。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /username:( |)req.user.username/,
    'You should be passing the variable username with req.user.username into the render function of the profile page'
  );
```

