---
id: 58966a17f9fc0f352b528e6d
title: 新使用者註冊
challengeType: 2
forumTopicId: 301561
dashedName: registration-of-new-users
---

# --description--

現在你需要允許新使用者在你的網站上註冊帳號。在首頁的 `res.render` 中，為傳遞的物件添加一個新的變數 - `showRegistration: true`。當你重新整理頁面後，你應該會看到已經在你的 `index.pug` 檔案中建立的註冊表單。此表單設定為對 `/register` 進行 **POST**，所以請建立該路由，並依照以下邏輯將使用者物件新增到資料庫中。

註冊路由的邏輯應該如下：

1. 註冊新的使用者
2. 驗證新的使用者
3. 重新導向到 `/profile`

步驟 1 的邏輯應該如下：

1. 使用 `findOne` 查詢資料庫
2. 如果有錯誤，請使用該錯誤呼叫 `next`。
3. 如果傳回了使用者，則重新導向回首頁
4. 如果找不到使用者且沒有發生錯誤，則使用使用者名稱和密碼對資料庫執行 `insertOne`。只要那裡沒有發生錯誤，就呼叫 `next` 以進入步驟 2，驗證新使用者，你已經在 `POST /login` 路由中撰寫了該邏輯。

```js
app.route('/register')
  .post((req, res, next) => {
    myDataBase.findOne({ username: req.body.username }, (err, user) => {
      if (err) {
        next(err);
      } else if (user) {
        res.redirect('/');
      } else {
        myDataBase.insertOne({
          username: req.body.username,
          password: req.body.password
        },
          (err, doc) => {
            if (err) {
              res.redirect('/');
            } else {
              // The inserted document is held within
              // the ops property of the doc
              next(null, doc.ops[0]);
            }
          }
        )
      }
    })
  },
    passport.authenticate('local', { failureRedirect: '/' }),
    (req, res, next) => {
      res.redirect('/profile');
    }
  );
```

當你認為頁面已經正確時，請提交你的頁面。如果你遇到錯誤，你可以<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#registration-of-new-users-11" target="_blank" rel="noopener noreferrer nofollow">查看目前為止完成的專案</a>。

**注意：** 從此點開始，可能會出現與使用 *picture-in-picture* 瀏覽器相關的問題。如果你正在使用提供在編輯器中預覽應用程式的線上整合開發環境，建議將此預覽在新分頁中開啟。

# --hints--

你應該有一個 `/register` 路由並在首頁顯示一個註冊表單。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /showRegistration:( |)true/gi,
    'You should be passing the variable showRegistration as true to your render function for the homepage'
  );
  assert.match(
    data,
    /register[^]*post[^]*findOne[^]*username:( |)req.body.username/gi,
    'You should have a route that accepts a POST request on /register that queries the db with findOne and the query being username: req.body.username'
  );
```

註冊應該可行。

```js
  const url = code;
  const user = `freeCodeCampTester${Date.now()}`;
  const xhttp = new XMLHttpRequest();
  xhttp.onreadystatechange = function () {
    if (this.readyState == 4 && this.status == 200) {
      test(this);
    } else {
      throw new Error(`${this.status} ${this.statusText}`);
    }
  };
  xhttp.open('POST', url + '/register', true);
  xhttp.setRequestHeader('Content-type', 'application/x-www-form-urlencoded');
  xhttp.send(`username=${user}&password=${user}`);
  function test(xhttpRes) {
    const data = xhttpRes.responseText;
    assert.match(
      data,
      /Profile/gi,
      'Register should work, and redirect successfully to the profile.'
    );
  }
```

登入應該能正常運作。

```js
  const url = code;
  const user = `freeCodeCampTester${Date.now()}`;
  const xhttpReg = new XMLHttpRequest();
  xhttpReg.onreadystatechange = function () {
    if (this.readyState == 4 && this.status == 200) {
      login();
    } else {
      throw new Error(`${this.status} ${this.statusText}`);
    }
  };
  xhttpReg.open('POST', url + '/register', true);
  xhttpReg.setRequestHeader(
    'Content-type',
    'application/x-www-form-urlencoded'
  );
  xhttpReg.send(`username=${user}&password=${user}`);
  function login() {
    const xhttp = new XMLHttpRequest();
    xhttp.onreadystatechange = function () {
      if (this.readyState == 4 && this.status == 200) {
        test(this);
      } else {
        throw new Error(`${this.status} ${this.statusText}`);
      }
    };
    xhttp.open('POST', url + '/login', true);
    xhttp.setRequestHeader('Content-type', 'application/x-www-form-urlencoded');
    xhttp.send(`username=${user}&password=${user}`);
  }
  function test(xhttpRes) {
    const data = xhttpRes.responseText;
    assert.match(
      data,
      /Profile/gi,
      'Login should work if previous test was done successfully and redirect successfully to the profile.'
    );
    assert.match(
      data,
      new RegExp(user, 'g'),
      'The profile should properly display the welcome to the user logged in'
    );
  }
```

登出應該能正常運作。

```js
  $.ajax({
    url: code + '/logout',
    type: 'GET',
    xhrFields: { withCredentials: true }
  }).then(
    (data) => {
      assert.match(data, /Home/gi, 'Logout should redirect to home');
    },
    (xhr) => {
      throw new Error(xhr.statusText);
    }
  );
```

登出後，`profile` 不應再運作。

```js
  $.ajax({
    url: code + '/profile',
    type: 'GET',
    crossDomain: true,
    xhrFields: { withCredentials: true }
  }).then(
    (data) => {
      assert.match(
        data,
        /Home/gi,
        'Profile should redirect to home when we are logged out now again'
      );
    },
    (xhr) => {
      throw new Error(xhr.statusText);
    }
  );
```

