---
id: 589690e6f9fc0f352b528e6e
title: 使用模組清理你的專案
challengeType: 2
forumTopicId: 301549
dashedName: clean-up-your-project-with-modules
---

# --description--

目前，你擁有的所有內容都在你的 `server.js` 檔案中。這可能導致難以管理且不易擴充的程式碼。創建 2 個新檔案：`routes.js` 和 `auth.js`。

兩者都應該以以下程式碼開始：

```js
module.exports = function (app, myDataBase) {

}
```

現在，在你的伺服端檔案頂部，像這樣引入這些檔案：`const routes = require('./routes.js');` 在你成功建立與資料庫的連線後，像這樣實例化它們：`routes(app, myDataBase)`

最後，將你伺服端中的所有路由貼到你的新檔案中，並從你的伺服端檔案中移除它們。也將 `ensureAuthenticated` 函式帶過來，因為它是專門為路由創建的。現在，你必須正確地在 `routes.js` 檔案中匯出行之前的最上方，添加所使用的相依性，例如 `const passport = require('passport');`。

持續添加它們直到不再有錯誤，且你的伺服端檔案不再有任何路由（**除了在 catch 區塊中的路由**）！

在你的 `auth.js` 檔案中對所有與驗證相關的事項執行相同的操作，例如序列化和設定區域的策略，並從你的伺服端檔案中刪除它們。務必將相依性添加進去，並在伺服端的相同位置呼叫 `auth(app, myDataBase)`。

當你認為已經正確完成時，提交你的頁面。如果你遇到錯誤，你可以<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#clean-up-your-project-with-modules-2" target="_blank" rel="noopener noreferrer nofollow">查看完成專案的範例</a>。

# --hints--

模組應該存在。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require\s*\(('|")\.\/routes(\.js)?\1\)/gi,
    'You should have required your new files'
  );
  assert.match(
    data,
    /client\s*\.db[^]*routes/gi,
    'Your new modules should be called after your connection to the database'
  );
```

