---
id: 58965611f9fc0f352b528e6c
title: 登出使用者
challengeType: 2
forumTopicId: 301560
dashedName: logging-a-user-out
---

# --description--

建立登出邏輯很簡單。路由只需取消使用者的驗證，並重新導向到首頁，而不是呈現任何視圖。

在 passport 中，取消驗證使用者就像在重新導向之前呼叫 `req.logout()` 一樣簡單。為了做到這點，將此 `/logout` 路由添加進去：

```js
app.route('/logout')
  .get((req, res) => {
    req.logout();
    res.redirect('/');
});
```

你可能已經注意到你沒有控制遺失頁面（404）。在 Node 中處理這個問題的常見方法是使用以下中介層。請在所有其他路由之後將它添加進去：

```js
app.use((req, res, next) => {
  res.status(404)
    .type('text')
    .send('Not Found');
});
```

當你認為頁面已經正確時，請提交你的頁面。如果你遇到錯誤，你可以<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#logging-a-user-out-10" target="_blank" rel="noopener noreferrer nofollow">查看至此為止完成的專案</a>。

# --hints--

應該在你的 `/logout` 路由中呼叫 `req.logout()`。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /req.logout/gi,
    'You should be calling req.logout() in your /logout route'
  );
```

`/logout` 應該重新導向到首頁。

```js
  const url = new URL("/logout", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Home page/gi,
    'When a user logs out they should be redirected to the homepage'
  );
```

