---
id: 589a69f5f9fc0f352b528e71
title: 社交驗證 II 的實作
challengeType: 2
forumTopicId: 301557
dashedName: implementation-of-social-authentication-ii
---

# --description--

設定你的 GitHub 認證的最後一個部分是創建策略本身。`passport-github@~1.1.0` 已經被添加為相依性，所以在你的 `auth.js` 檔案中以 `GitHubStrategy` 這樣 require 它：`const GitHubStrategy = require('passport-github').Strategy;`。別忘了 require 並組態 `dotenv` 以使用你的環境變數。

要設定 GitHub 策略，你必須告訴 Passport 使用具現化的 `GitHubStrategy`，它接受 2 個引數：一個物件（包含 `clientID`、`clientSecret` 和 `callbackURL`）以及一個在使用者成功驗證時會被呼叫的函式，該函式會判斷使用者是否為新用戶，以及最初要在使用者的資料庫物件中儲存哪些欄位。這在許多策略中很常見，但某些策略可能需要更多資訊，如該特定策略的 GitHub README 所述。例如，Google 也需要一個 *scope*，它決定你的請求要求傳回哪種類型的資訊，並要求使用者批准此類存取。

你目前實作的策略是使用 GitHub 帳號和 OAuth 2.0 語彙單元來驗證使用者。建立應用程式時取得的客戶端 ID 和密鑰會作為選項提供給策略。該策略還需要一個 `verify` 回呼，該回呼會接收存取語彙單元和選用的重新整理語彙單元，以及包含已驗證使用者 GitHub 評測的 `profile`。`verify` 回呼必須呼叫 `cb` 並提供使用者以完成驗證。

以下是你目前的新策略應該呈現的樣子：

```js
passport.use(new GitHubStrategy({
  clientID: process.env.GITHUB_CLIENT_ID,
  clientSecret: process.env.GITHUB_CLIENT_SECRET,
  callbackURL: /*INSERT CALLBACK URL ENTERED INTO GITHUB HERE*/
},
  function(accessToken, refreshToken, profile, cb) {
    console.log(profile);
    //Database logic here with callback containing your user object
  }
));
```

你的驗證尚未成功，且如果沒有資料庫邏輯和回呼，它實際上會發出錯誤，但如果你嘗試，它應該會將你的 GitHub 評測記錄到你的主控台！

當你認為已經完成正確時，提交你的頁面。如果你遇到錯誤，你可以<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implementation-of-social-authentication-ii-4" target="_blank" rel="noopener noreferrer nofollow">查看至此為止完成的專案</a>。

# --hints--

應該為你的專案添加 `passport-github` 相依性。

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport-github',
    'Your project should list "passport-github" as a dependency'
  );
```

應該要 require `passport-github`。

```js
  const url = new URL("/_api/auth.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')passport-github("|')/gi,
    'You should have required passport-github'
  );
```

GitHub 策略到目前為止應該已正確設定。

```js
  const url = new URL("/_api/auth.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport\.use.*new GitHubStrategy/gis,
    'Passport should use a new GitHubStrategy'
  );
  assert.match(
    data,
    /callbackURL:\s*("|').*("|')/gi,
    'You should have a callbackURL'
  );
  assert.match(
    data,
    /process\.env(\.GITHUB_CLIENT_SECRET|\[(?<q>"|')GITHUB_CLIENT_SECRET\k<q>\])/g,
    'You should use process.env.GITHUB_CLIENT_SECRET'
  );
  assert.match(
    data,
    /process\.env(\.GITHUB_CLIENT_ID|\[(?<q>"|')GITHUB_CLIENT_ID\k<q>\])/g,
    'You should use process.env.GITHUB_CLIENT_ID'
  );
```
