---
id: 589fc830f9fc0f352b528e74
title: 設定環境
challengeType: 2
forumTopicId: 301566
dashedName: set-up-the-environment
---

# --description--

接下來的挑戰將會使用 `chat.pug` 檔案。所以，在你的 `routes.js` 檔案中，為 `/chat` 添加一個 GET 路由，該路由使用 `ensureAuthenticated`，並呈現 `chat.pug`，同時將 `{ user: req.user }` 作為引數傳遞給回應。現在，修改你現有的 `/auth/github/callback` 路由，設定 `req.session.user_id = req.user.id`，並重新導向到 `/chat`。

`socket.io@~2.3.0` 已經作為相依性添加，因此請在你的伺服器中使用 `http`（Nodejs 內建）如下要求/實例化它：

```javascript
const http = require('http').createServer(app);
const io = require('socket.io')(http);
```

既然 *http* 伺服器已掛載在 *express app* 上，你需要從 *http* 伺服器監聽。將包含 `app.listen` 的那一行改為 `http.listen`。

首先需要控制代碼的是監聽來自客戶端的新連接。<dfn>on</dfn> 關鍵字正是用來做這件事－監聽特定的事件。它需要 2 個引數：一個包含被觸發事件標題的字串，以及一個用來傳遞資料的函式（程式）。在我們的連接監聽器案例中，使用 `socket` 來定義第二個引數中的資料。`socket` 是一個已連接的個別客戶端。

要監聽對你的伺服器的連接，請在你的資料庫連接中添加以下內容：

```javascript
io.on('connection', socket => {
  console.log('A user has connected');
});
```

現在，為了讓客戶端連線，你只需要將以下內容添加到經過驗證後由頁面載入的 `client.js` 中：

```js
/*global io*/
let socket = io();
```

該註解抑制了你通常會看到的錯誤，因為在該檔案中未定義 `io`。你已經在 `chat.pug` 頁面中為 Socket.IO 函式庫添加了可靠的 CDN。

現在嘗試載入你的應用程式並進行驗證，你應該會在你的伺服端主控台看到 `A user has connected`。

**注意：**`io()` 僅在連接到同一 url/伺服器上託管的 socket 時有效。若要連接到其他地方託管的外部 socket，則會使用 `io.connect('URL');`。

當你認為已經完成頁面時，請提交。如果你遇到錯誤，你可以<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-the-environment-6" target="_blank" rel="noopener noreferrer nofollow">查看至此為止完成的專案</a>。

# --hints--

`socket.io` 應該是一個相依性。

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'socket.io',
    'Your project should list "socket.io" as a dependency'
  );
```

你應該正確地 require 並實例化 `http` 為 `http`。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /http.*=.*require.*('|")http\1/s,
    'Your project should list "http" as a dependency'
  );
```

你應該正確地 require 並實例化 `socket.io` 為 `io`。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io.*=.*require.*('|")socket.io\1.*http/s,
    'You should correctly require and instantiate socket.io as io.'
  );
```

Socket.IO 應該正在監聽連接。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io.on.*('|")connection\1.*socket/s,
    'io should listen for "connection" and socket should be the 2nd arguments variable'
  );
```

你的客戶端應該連接到你的伺服器。

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.*=.*io/s,
    'Your client should be connection to server with the connection defined as socket'
  );
```

