---
id: 589fc832f9fc0f352b528e78
title: 宣布新使用者
challengeType: 2
forumTopicId: 301546
dashedName: announce-new-users
---

# --description--

許多聊天室能夠在使用者連線或斷線時公告，並將該訊息顯示給所有連線中的使用者。既然你已經在連線和斷線時發出 `event`，你只需修改此 `event` 以支援此特性。最合邏輯的做法是隨 `event` 傳送 3 筆資料：連線/斷線使用者的使用者名稱、目前的使用者數量，以及該使用者名稱是連線還是斷線。

將事件名稱改為 `'user'`，並傳遞一個包含欄位 `username`、`currentUsers` 和 `connected` 的物件（若為使用者連線則為 `true`，若為斷線則為 `false`）。務必同時更改 `'user count'` 事件，並將斷線事件的 `connected` 欄位設為傳送 `false`，而非連線事件發出的 `true`。

```js
io.emit('user', {
  username: socket.request.user.username,
  currentUsers,
  connected: true
});
```

現在你的客戶端將擁有所有必要的資訊，以正確顯示目前的使用者數量並在使用者連線或斷線時進行公告！要在客戶端控制這個事件，我們應該監聽 `'user'`，然後使用 jQuery 將 `#num-users` 的文字更新為 `'{NUMBER} users online'`，並且將 `<li>` 新增到 id 為 `messages` 的無序列表中，內容為 `'{NAME} has {joined/left} the chat.'`。

這個的實作可能看起來如下：

```js
socket.on('user', data => {
  $('#num-users').text(data.currentUsers + ' users online');
  let message =
    data.username +
    (data.connected ? ' has joined the chat.' : ' has left the chat.');
  $('#messages').append($('<li>').html('<b>' + message + '</b>'));
});
```

當你認為已經完成頁面時，請提交你的頁面。如果你遇到錯誤，可以查看<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135/3#announce-new-users-10" target="_blank" rel="noopener noreferrer nofollow">到目前為止完成的專案</a>。

# --hints--

應該發出帶有 `username`、`currentUsers` 和 `connected` 的 `'user'` 事件。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  // Regex is lenient to match both `username` and `name` as the key on purpose.
  assert.match(
    data,
    /io.emit.*('|")user\1.*name.*currentUsers.*connected/s,
    'You should have an event emitted named user sending name, currentUsers, and connected'
  );
```

客戶端應該正確控制代碼並顯示來自 `user` 事件的新資料。

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")user\1[^]*num-users/s,
    'You should change the text of "#num-users" within on your client within the "user" event listener to show the current users connected'
  );
  assert.match(
    data,
    /socket.on.*('|")user\1[^]*messages.*li/s,
    'You should append a list item to "#messages" on your client within the "user" event listener to announce a user came or went'
  );
```

