---
id: 589fc832f9fc0f352b528e79
title: 傳送並顯示聊天訊息
challengeType: 2
forumTopicId: 301562
dashedName: send-and-display-chat-messages
---

# --description--

是時候開始允許客戶端傳送聊天訊息到伺服器，再由伺服器發送給所有客戶端了！在你的 `client.js` 檔案中，你應該已經看到有一個區塊程式碼在處理表單送出時的情況。

```js
$('form').submit(function() {
  /*logic*/
});
```

在表單提交的程式碼中，你應該在定義 `messageToSend` 之後但在清除文字方塊 `#m` 之前發出一個事件。該事件應命名為 `'chat message'`，資料應該只是 `messageToSend`。

```js
socket.emit('chat message', messageToSend);
```

現在，在你的伺服器上，你應該監聽 `socket` 的 `'chat message'` 事件，資料名稱為 `message`。一旦收到該事件，應該使用 `io.emit` 向所有 `socket` 發出 `'chat message'` 事件，並傳送包含 `username` 和 `message` 的資料物件。

在 `client.js` 中，你現在應該監聽 `event` `'chat message'`，並在收到時，將包含使用者名稱、冒號和訊息的列表項目附加到 `#messages`！

此時，聊天應該已完全運作並在所有客戶端之間傳送訊息！

當你認為頁面已正確時，請提交你的頁面。如果你遇到錯誤，你可以<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#send-and-display-chat-messages-11" target="_blank" rel="noopener noreferrer nofollow">查看至此為止完成的專案</a>。

# --hints--

伺服器應該監聽 `'chat message'`，然後正確地發送它。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")chat message('|")[^]*io.emit.*('|")chat message('|").*username.*message/s,
    'Your server should listen to the socket for "chat message" then emit to all users "chat message" with name and message in the data object'
  );
```

客戶端應該正確控制代碼並顯示來自事件 `'chat message'` 的新資料。

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")chat message('|")[^]*messages.*li/s,
    'You should append a list item to #messages on your client within the "chat message" event listener to display the new message'
  );
```

