---
id: 589fc831f9fc0f352b528e77
title: 使用 Socket.IO 進行身份驗證
challengeType: 2
forumTopicId: 301548
dashedName: authentication-with-socket-io
---

# --description--

目前，你無法判斷誰連接到你的 web socket。雖然 `req.user` 包含使用者物件，但那僅在你的使用者與網頁伺服器交談時有效，而使用 web socket 時你沒有 `req`（請求），因此也沒有使用者資料。解決知道誰連接到你的 web socket 問題的一種方法是解析並解碼包含 passport session 的 cookie，然後反序列化以取得使用者物件。幸運的是，NPM 上有一個軟體包專門用於此，將曾經複雜的任務變得簡單！

`passport.socketio@~3.7.0`、`connect-mongo@~3.2.0` 和 `cookie-parser@~1.4.5` 已經被添加為相依套件。分別以 `passportSocketIo`、`MongoStore` 和 `cookieParser` 引入它們。此外，我們需要從先前引入的 `express-session` 初始化一個新的記憶體儲存。它應該看起來像這樣：

```js
const MongoStore = require('connect-mongo')(session);
const URI = process.env.MONGO_URI;
const store = new MongoStore({ url: URI });
```

現在我們只需要告訴 Socket.IO 使用它並設定選項。請確定這是在現有的 socket 程式碼之前加入，而不是在現有的連接監聽器中。對你的伺服端來說，應該看起來像這樣：

```js
io.use(
  passportSocketIo.authorize({
    cookieParser: cookieParser,
    key: 'express.sid',
    secret: process.env.SESSION_SECRET,
    store: store,
    success: onAuthorizeSuccess,
    fail: onAuthorizeFail
  })
);
```

請注意，為 Socket.IO 組態 Passport 認證與我們為 API 組態 `session` 中介層的方式非常相似。這是因為它們旨在使用相同的認證方法——從 cookie 取得 session id 並進行驗證。

之前，當我們組態 `session` 中介層時，我們沒有明確設定用於 session 的 cookie 名稱（`key`）。這是因為 `session` 軟體包使用了預設值。現在我們已經添加了另一個需要存取 cookie 中相同值的軟體包，我們需要在兩個設定物件中明確設定 `key` 值。

請務必將帶有 cookie 名稱的 `key` 添加到與 Socket.IO `key` 相符的 `session` 中介層中。此外，請在設定 `saveUninitialized: true` 附近，將 `store` 引用添加到選項中。這是告訴 Socket.IO 要關聯哪個 session 所必須的。

<hr>

現在，定義 `success` 和 `fail` 回呼函式。

```js
function onAuthorizeSuccess(data, accept) {
  console.log('successful connection to socket.io');

  accept(null, true);
}

function onAuthorizeFail(data, message, error, accept) {
  if (error) throw new Error(message);
  console.log('failed connection to socket.io:', message);
  accept(null, false);
}
```

使用者物件現在可以在你的 socket 物件中以 `socket.request.user` 方式存取。 例如，現在你可以添加以下內容：

```js
console.log('user ' + socket.request.user.username + ' connected');
```

它會在伺服器主控台記錄誰已連線！

當你認為已經正確完成時，提交你的頁面。如果你遇到錯誤，你可以 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#authentication-with-socketio-9" target="_blank" rel="noopener noreferrer nofollow">查看目前為止的專案內容</a>。

# --hints--

`passport.socketio` 應該是一個相依性。

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport.socketio',
    'Your project should list "passport.socketio" as a dependency'
  );
```

`cookie-parser` 應該是一個相依性。

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'cookie-parser',
    'Your project should list "cookie-parser" as a dependency'
  );
```

passportSocketIo 應該被正確 require。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require\((['"])passport\.socketio\1\)/gi,
    'You should correctly require and instantiate "passport.socketio"'
  );
```

passportSocketIo 應該正確設定。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io\.use\(\s*\w+\.authorize\(/,
    'You should register "passport.socketio" as socket.io middleware and provide it correct options'
  );
```

