---
id: 66db529d37ad966480ebb633
title: Costruisci un programma per scegliere il pranzo
challengeType: 26
dashedName: build-a-lunch-picker-program
---

# --description--

In questo laboratorio costruirai un programma che aiuta a gestire le opzioni per il pranzo. Lavorerai con un array di pranzi, aggiungerai e rimuoverai elementi dall'array e selezionerai casualmente un'opzione per il pranzo.

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti creare una variabile `lunches` e assegnarle un array vuoto che sarà usato per archiviare gli elementi del pranzo.

2. Dovresti creare una funzione `addLunchToEnd` che prende un array come primo argomento e una stringa come secondo argomento. La funzione dovrebbe:
    - Aggiungere la stringa alla fine dell'array.
    - Stampare nella console la stringa `[Lunch Item] added to the end of the lunch menu.`, dove `[Lunch Item]` è la stringa passata alla funzione.
    - Restituire l'array aggiornato.

3. Dovresti creare una funzione `addLunchToStart` che prende un array come primo argomento e una stringa come secondo argomento. La funzione dovrebbe:
    - Aggiungere la stringa all'inizio dell'array.
    - Stampare nella console la stringa `[Lunch Item] added to the start of the lunch menu.`, dove `[Lunch Item]` è la stringa passata alla funzione.
    - Restituire l'array aggiornato.

4. Dovresti creare una funzione `removeLastLunch` che prende un array come argomento. La funzione dovrebbe:
   - Rimuovere l'ultimo elemento dall'array.
   - Se la rimozione ha successo, stampare nella console la stringa `[Lunch Item] removed from the end of the lunch menu.`, dove `[Lunch Item]` è l'elemento rimosso dall'array.
   - Se l'array è vuoto, stampare nella console la stringa `"No lunches to remove."`.
   - Restituire l'array aggiornato.

5. Dovresti creare una funzione `removeFirstLunch` che prende un array come argomento. La funzione dovrebbe:
   - Rimuovere il primo elemento dall'array.
   - Se la rimozione ha successo, stampare nella console la stringa `[Lunch Item] removed from the start of the lunch menu.`, dove `[Lunch Item]` è l'elemento rimosso dall'array.
   - Se l'array è vuoto, stampare nella console la stringa `"No lunches to remove."`.
   - Restituire l'array aggiornato.

6. Dovresti creare una funzione `getRandomLunch` che prende un array come argomento. La funzione dovrebbe:
   - Selezionare un elemento casuale dall'array.
   - Se ha successo, stampare nella console la stringa `Randomly selected lunch: [Lunch Item]`, dove `[Lunch Item]` è un elemento casuale dell'array.
   - Se l'array è vuoto, stampare nella console la stringa `"No lunches available."`.

7. Dovresti creare una funzione `showLunchMenu` che prende un array come argomento e:
   - Se ci sono elementi nell'array, stampa nella console la stringa `Menu items: [Lunch Item], [Lunch Item]...`, dove ogni `[Lunch item]` è uno degli elementi dell'array, in ordine.
   - Se l'array è vuoto, stampa nella console la stringa `"The menu is empty."`.

# --hints--

Dovresti dichiarare una variabile `lunches` e assegnarle un array vuoto per archiviare gli elementi del pranzo.

```js
assert.isArray(lunches);
```

Dovresti definire una funzione `addLunchToEnd`.

```js
assert.isFunction(addLunchToEnd);
```

La funzione `addLunchToEnd` dovrebbe avere due parametri.

```js
assert.lengthOf(addLunchToEnd, 2);
```

`addLunchToEnd(lunches, "Tacos")` dovrebbe stampare nella console la stringa `"Tacos added to the end of the lunch menu."`.

```js
const tempArr = [];
const temp = console.log;
const testLunches = []
try {
  console.log = obj => tempArr.push(obj);
  addLunchToEnd(testLunches, "Tacos");
  assert.strictEqual(tempArr[0], "Tacos added to the end of the lunch menu.");
  addLunchToEnd(testLunches, "Pizza");
  assert.strictEqual(tempArr[1], "Pizza added to the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToEnd(["Pizza", "Tacos"], "Burger")` dovrebbe restituire `["Pizza", "Tacos", "Burger"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToEnd(["Pizza", "Tacos"], "Burger"), ["Pizza", "Tacos", "Burger"]);
  assert.deepStrictEqual(addLunchToEnd(["Fries"], "Tacos"), ["Fries", "Tacos"]);
} finally {
  console.log = temp;
}
```

Dovresti definire una funzione `addLunchToStart`.

```js
assert.isFunction(addLunchToStart);
```

La funzione `addLunchToStart` dovrebbe avere due parametri.

```js
assert.lengthOf(addLunchToStart, 2);
```

`addLunchToStart(lunches, "Sushi")` dovrebbe stampare nella console la stringa `"Sushi added to the start of the lunch menu."`.

```js
const tempArr = [];
const testLunches = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  addLunchToStart(testLunches, "Sushi");
  assert.strictEqual(tempArr[0], "Sushi added to the start of the lunch menu.");  
  addLunchToStart(testLunches, "Burger");
  assert.strictEqual(tempArr[1], "Burger added to the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToStart(["Burger", "Sushi"], "Pizza")` dovrebbe restituire `["Pizza", "Burger", "Sushi"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToStart(["Burger", "Sushi"], "Pizza"), ["Pizza", "Burger", "Sushi"]);
  assert.deepStrictEqual(addLunchToStart(["Noodles"], "Sushi"), ["Sushi", "Noodles"]);
} finally {
  console.log = temp;
}
```

Dovresti definire una funzione `removeLastLunch`.

```js
assert.isFunction(removeLastLunch)
```

La funzione `removeLastLunch` dovrebbe avere un parametro.

```js
assert.lengthOf(removeLastLunch, 1);
```

Quando l'array di input è vuoto, la funzione `removeLastLunch` dovrebbe stampare nella console la stringa `"No lunches to remove."`.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Stew", "Soup", "Toast"])` dovrebbe stampare nella console la stringa `"Toast removed from the end of the lunch menu."`.

```js
let testLunches = ["Stew", "Soup", "Toast"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "Toast removed from the end of the lunch menu.");
  testLunches = ["Rice", "Pizza"];
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[1], "Pizza removed from the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Sushi", "Pizza", "Noodles"])` dovrebbe restituire `["Sushi", "Pizza"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeLastLunch(["Sushi", "Pizza", "Noodles"]), ["Sushi", "Pizza"]);
  assert.deepStrictEqual(removeLastLunch(["Pizza", "Burger"]), ["Pizza"]);
} finally {
  console.log = temp;
}
```

Dovresti definire una funzione `removeFirstLunch`.

```js
assert.isFunction(removeFirstLunch);
```

La funzione `removeFirstLunch` dovrebbe avere un solo parametro.

```js
assert.lengthOf(removeFirstLunch, 1);
```

Quando l'array di input è vuoto, la funzione `removeFirstLunch` dovrebbe stampare nella console la stringa `"No lunches to remove."`.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Salad", "Eggs", "Cheese"])` dovrebbe stampare nella console la stringa `"Salad removed from the start of the lunch menu."`.

```js
let testLunches = ["Salad", "Eggs", "Cheese"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "Salad removed from the start of the lunch menu.");
  testLunches = ["Stew", "Soup", "Toast"]
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[1], "Stew removed from the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Sushi", "Pizza", "Burger"])` dovrebbe restituire `["Pizza", "Burger"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeFirstLunch(["Sushi", "Pizza", "Burger"]), ["Pizza", "Burger"]);
  assert.deepStrictEqual(removeFirstLunch(["Pizza", "Burger"]), ["Burger"]);
} finally {
  console.log = temp;
}
```

`addLunchToEnd`, `addLunchToStart`, `removeLastLunch` e `removeFirstLunch` dovrebbero restituire lo stesso array passato come argomento dopo averlo aggiornato.

```js
const temp = console.log;
console.log = () => {}
const testLunches = ["Spam"];
try {
  assert.equal(addLunchToEnd(testLunches, "Pizza"), testLunches);
  assert.equal(addLunchToStart(testLunches, "Caviar"), testLunches);
  assert.equal(removeLastLunch(testLunches), testLunches);
  assert.equal(removeFirstLunch(testLunches), testLunches);
} finally {
  console.log = temp;
}
```

Dovresti definire una funzione `getRandomLunch`.

```js
assert.isFunction(getRandomLunch);
```

La funzione `getRandomLunch` dovrebbe avere un solo parametro.

```js
assert.lengthOf(getRandomLunch, 1);
```

Quando l'array di input è vuoto, la funzione `getRandomLunch` dovrebbe stampare nella console la stringa `"No lunches available."`.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  getRandomLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches available.");
} finally {
  console.log = temp;
}
```

Quando l'array di input non è vuoto, la funzione `getRandomLunch` dovrebbe stampare una stringa nel formato `Randomly selected lunch: [Lunch Item]` nella console.

```js
const testLunches = ["Fish", "Fries", "Roast"];
const tempRandom = Math.random;
const tempArr = [];
const temp = console.log;

try {
  console.log = obj => tempArr.push(obj);
  // check that it correctly outputs the first item
  Math.random = () => 0;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[0], `Randomly selected lunch: ${testLunches[0]}`);

  // Second item
  Math.random = () => 0.4;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[1], `Randomly selected lunch: ${testLunches[1]}`);

  // Third item
  Math.random = () => 0.8;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[2], `Randomly selected lunch: ${testLunches[2]}`);

} finally {
  Math.random = tempRandom;
  console.log = temp;
}
```

La funzione `getRandomLunch` non dovrebbe modificare l'array passato come argomento.

```js
const testLunches = ["Fish", "Fries", "Roast"];
let originalLunches = [...testLunches];
const tempRandom = Math.random;

try {
  Math.random = () => 0;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.4;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.8;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  const emptyLunches = [];
  let originalEmpty = [...emptyLunches];
  getRandomLunch(emptyLunches);
  assert.sameOrderedMembers(emptyLunches, originalEmpty);

} finally {
  Math.random = tempRandom;
}
```

Dovresti definire una funzione `showLunchMenu`.

```js
assert.isFunction(showLunchMenu);
```

La funzione `showLunchMenu` dovrebbe avere un solo parametro.

```js
assert.lengthOf(showLunchMenu, 1);
```

Quando l'array di input è vuoto, la funzione `showLunchMenu` dovrebbe stampare nella console la stringa `"The menu is empty."`.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], "The menu is empty.");
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Greens", "Corns", "Beans"])` dovrebbe stampare `"Menu items: Greens, Corns, Beans"` nella console.

```js
const testLunches = ["Greens", "Corns", "Beans"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Pizza", "Burger", "Fries", "Salad"])` dovrebbe stampare `"Menu items: Pizza, Burger, Fries, Salad"` nella console.

```js
const testLunches = ["Pizza", "Burger", "Fries", "Salad"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const lunches = [];

// Add a new lunch to the end of the list
function addLunchToEnd(lunches, newLunch) {
  lunches.push(newLunch);
  console.log(`${newLunch} added to the end of the lunch menu.`);
  return lunches;
}

// Add a new lunch to the start of the list
function addLunchToStart(lunches, newLunch) {
  lunches.unshift(newLunch);
  console.log(`${newLunch} added to the start of the lunch menu.`);
  return lunches;
}

// Remove the last lunch from the list
function removeLastLunch(lunches) {
  const removedLunch = lunches.pop();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the end of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Remove the first lunch from the list
function removeFirstLunch(lunches) {
  const removedLunch = lunches.shift();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the start of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Function to get a random lunch from the list
function getRandomLunch(lunches) {
  if (lunches.length === 0) {
    console.log("No lunches available.");
  }
  else {
    const randomIndex = Math.floor(Math.random() * lunches.length);
    const randomLunch = lunches[randomIndex];
    console.log(`Randomly selected lunch: ${randomLunch}`);
  }
}

// Function to display all the lunches in the list
function showLunchMenu(lunches) {
  if (lunches.length === 0) {
    console.log("The menu is empty.");
  }
  else {
    console.log(`Menu items: ${lunches.join(', ')}`);
  }
}

addLunchToEnd(lunches, "Tacos");
addLunchToStart(lunches, "Sushi");
removeLastLunch(lunches);
removeFirstLunch(lunches);
getRandomLunch(lunches);
showLunchMenu(lunches);

```
