---
id: 668f08ea07b99b1f4a91acab
title: Crea una pagina di ricetta
challengeType: 25
dashedName: build-a-recipe-page
demoType: onClick
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una dichiarazione `!DOCTYPE html`.
1. Dovresti avere un elemento `html` con `lang` impostato su `en`.
1. Dovresti avere un elemento `head` contenente un elemento `title` con il nome della tua ricetta e un elemento `meta` con un attributo `charset` impostato su `UTF-8`.
1. Dovresti avere un elemento `body`.
1. Dovresti avere un elemento `h1` con il nome della tua ricetta.
1. Dovresti avere un elemento `p` che introduce la ricetta sotto `h1`.
1. Dovresti avere un elemento `h2` con il testo `Ingredients` per la sezione degli ingredienti.
1. Dovresti avere una lista non ordinata (elemento `ul`) con almeno quattro elementi di lista (elementi `li`) che elenca i tuoi ingredienti sotto il primo elemento `h2`.
1. Dovresti avere un secondo elemento `h2` con il testo `Instructions` per la sezione delle istruzioni.
1. Dovresti avere una lista ordinata (elemento `ol`) con almeno quattro elementi di lista che elenca i passaggi della ricetta in ordine, sotto il secondo `h2`.
1. Dovresti avere almeno un elemento `img` con un attributo `src` impostato su un'immagine valida, puoi usare `https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg` se vuoi, e un attributo `alt` che descrive l'immagine.

# --hints--

La tua pagina di ricetta dovrebbe avere una dichiarazione `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Dovresti avere un elemento `html` con `lang` impostato su `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/gi);
```

Dovresti avere un elemento `head` all'interno dell'elemento `html`.

```js
assert.match(code, /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i);
```

Dovresti avere un elemento `title` all'interno del tuo elemento `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i);
```

Il tuo elemento `title` dovrebbe avere il titolo della tua ricetta.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Dovresti avere un elemento `meta` all'interno del tuo elemento `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*meta[\s\S]*>[\s\S]*<\/\s*head\s*>/i);
```

Il tuo elemento `meta` dovrebbe avere il suo attributo `charset` impostato su `UTF-8`.

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

Dovresti avere un elemento `body` all'interno del tuo elemento `html`.

```js
assert.match(code, /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i);
```

Dovresti avere un elemento `h1` con il nome della tua ricetta.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Dovresti avere solo un elemento `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Dovresti avere un elemento `p` sotto il tuo elemento `h1`.

```js
assert.strictEqual(document.querySelector('h1')?.nextElementSibling, document.querySelector('p'));
```

Il tuo primo elemento `p` dovrebbe descrivere la tua ricetta.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Il tuo primo elemento `h2` dovrebbe avere il testo `Ingredients`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Ingredients');
```

Dovresti avere un elemento lista non ordinata sotto il tuo primo elemento `h2`.

```js
assert.strictEqual(document.querySelector('ul')?.previousElementSibling.tagName, 'H2');
```

Dovresti avere almeno quattro elementi di lista nella tua lista non ordinata con gli ingredienti.

```js
const els = document.querySelectorAll('ul > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

Il tuo secondo elemento `h2` dovrebbe avere il testo `Instructions`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Instructions');
```

Dovresti avere un elemento lista ordinata sotto il tuo secondo elemento `h2`.

```js
assert.strictEqual(document.querySelectorAll('h2')?.[1]?.nextElementSibling?.tagName, "OL");
```

Dovresti avere almeno quattro elementi di lista nella tua lista ordinata con le istruzioni.

```js
const els = document.querySelectorAll('ol > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

Dovresti avere almeno un elemento `img`.

```js
assert.exists(document.querySelector('img'));
```

Tutti i tuoi elementi `img` dovrebbero avere un attributo `src` valido e un valore.

```js
const img = document.querySelector('img');
const rawSrc = img?.getAttribute('src');
const resolvedSrc = img?.src;
const re = new RegExp(window.location.href, "ig");

assert.isAbove(rawSrc?.trim().length, 0, "The 'src' attribute must be explicitly set.");
assert.notMatch(resolvedSrc, re, "The 'src' should not start with the current page URL");

img.onload = () => {
  console.log('Image loaded successfully.');
};

img.onerror = (error) => {
  console.error('Image failed to load:', error);
  assert.fail("Your image's URL should be valid."); // Make the test instafail
};

if (img.complete) {
  img.onload && img.onload();
};
```

Tutti i tuoi elementi `img` dovrebbero avere un attributo `alt` per descrivere l'immagine.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Chocolate chip cookies recipe</title>
</head>

<body>
  <h1>Chocolate Chip Cookies</h1>
  <p>Welcome to the ultimate guide for making mini chocolate chip cookies! These bite-sized treats are perfect for
    satisfying your sweet tooth without overindulging. Follow this simple recipe to create delicious,
    crispy-on-the-outside, chewy-on-the-inside mini chocolate chip cookies that everyone will love.</p>
  <img src="https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg" alt="Ingredients for baking: three eggs, a bowl of flour, a glass of milk, and a whisk arranged on a wooden table.">
  <h2>Ingredients</h2>
  <ul>
    <li>1 cup all-purpose flour</li>
    <li>1/2 teaspoon baking soda</li>
    <li>1/4 cup unsalted butter, softened</li>
    <li>1/4 cup granulated sugar</li>
    <li>1/2 teaspoon vanilla extract</li>
    <li>1/2 cup mini chocolate chips</li>
  </ul>
  <h2>Instructions</h2>
  <ol>
    <li>Preheat your oven to 350°F (175°C) and line a baking sheet with parchment paper.</li>
    <li>In a bowl, whisk together the flour and baking soda.</li>
    <li>In another bowl, beat the butter, sugar, and vanilla extract until creamy.</li>
    <li>Gradually add the dry ingredients to the wet mixture, then fold in the mini chocolate chips.</li>
    <li>Drop small spoonfuls of dough onto the baking sheet.</li>
    <li>Bake for 8-10 minutes, then let cool before enjoying!</li>
  </ol>
</body>

</html>
```
