---
id: 587d7b8c367417b2b2512b56
title: "Export verwenden, um einen Codeblock zu teilen"
challengeType: 1
forumTopicId: 301219
dashedName: use-export-to-share-a-code-block
---

# --description--

Stellen Sie sich eine Datei namens `math_functions.js` vor, die mehrere Funktionen für mathematische Operationen enthält. Eine davon ist in einer Variablen `add` gespeichert, die zwei Zahlen entgegennimmt und deren Summe zurückgibt. Sie möchten diese Funktion in mehreren verschiedenen JavaScript-Dateien verwenden. Um sie mit diesen anderen Dateien zu teilen, müssen Sie sie zuerst `export`.

```js
export const add = (x, y) => {
  return x + y;
}
```

Das obige ist eine übliche Methode, um eine einzelne Funktion zu exportieren, aber Sie können dasselbe auch so erreichen:

```js
const add = (x, y) => {
  return x + y;
}

export { add };
```

Wenn Sie eine Variable oder Funktion exportieren, können Sie sie in einer anderen Datei importieren und verwenden, ohne den Code neu schreiben zu müssen. Sie können mehrere Dinge exportieren, indem Sie das erste Beispiel für jedes Element wiederholen, das Sie exportieren möchten, oder indem Sie sie alle in der Export-Anweisung des zweiten Beispiels zusammenfassen, so:

```js
export { add, subtract };
```

# --instructions--

Im Editor befinden sich zwei stringbezogene Funktionen. Exportieren Sie beide mit der Methode Ihrer Wahl.

# --hints--

Sie sollten `uppercaseString` korrekt exportieren.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(export\s+const\s+uppercaseString|export\s*{\s*(uppercaseString[^}]*|[^,]*,\s*uppercaseString\s*)})/g
  )
);
```

Sie sollten `lowercaseString` korrekt exportieren.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(export\s+const\s+lowercaseString|export\s*{\s*(lowercaseString[^}]*|[^,]*,\s*lowercaseString\s*)})/g
  )
);
```

# --seed--

## --seed-contents--

```js
const uppercaseString = (string) => {
  return string.toUpperCase();
}

const lowercaseString = (string) => {
  return string.toLowerCase()
}
```

# --solutions--

```js
export const uppercaseString = (string) => {
  return string.toUpperCase();
}

export const lowercaseString = (string) => {
  return string.toLowerCase()
}
```
