---
id: 587d7b8c367417b2b2512b54
title: "Getter und Setter verwenden, um den Zugriff auf ein Objekt zu steuern"
challengeType: 1
forumTopicId: 301220
dashedName: use-getters-and-setters-to-control-access-to-an-object
---

# --description--

Sie können Werte aus einem Objekt erhalten und den Wert einer Eigenschaft innerhalb eines Objekts festlegen.

Diese werden klassisch als <dfn>Getter</dfn> und <dfn>Setter</dfn> bezeichnet.

Getter-Funktionen sollen einfach den Wert einer privaten Variablen eines Objekts an den Benutzer zurückgeben (holen), ohne dass der Benutzer direkt auf die private Variable zugreift.

Setter-Funktionen sollen den Wert einer privaten Variablen eines Objekts basierend auf dem an die Setter-Funktion übergebenen Wert ändern (setzen). Diese Änderung kann Berechnungen beinhalten oder sogar den vorherigen Wert vollständig überschreiben.

```js
class Book {
  constructor(author) {
    this._author = author;
  }
  // getter
  get writer() {
    return this._author;
  }
  // setter
  set writer(updatedAuthor) {
    this._author = updatedAuthor;
  }
}
const novel = new Book('anonymous');
console.log(novel.writer);
novel.writer = 'newAuthor';
console.log(novel.writer);
```

Die Konsole würde die Strings `anonymous` und `newAuthor` anzeigen.

Beachten Sie die Syntax, die verwendet wird, um den Getter und Setter aufzurufen. Sie sehen nicht einmal wie Funktionen aus. Getter und Setter sind wichtig, weil sie interne Implementierungsdetails verbergen.

**Hinweis:** Es ist Konvention, den Namen einer privaten Variablen mit einem Unterstrich (`_`) zu beginnen. Diese Praxis macht eine Variable jedoch nicht automatisch privat.

# --instructions--

Verwenden Sie das Schlüsselwort `class`, um eine `Thermostat`-Klasse zu erstellen. Der `constructor` akzeptiert eine Fahrenheit-Temperatur.

Erstellen Sie in der Klasse einen `getter`, um die Temperatur in Celsius zu erhalten, und einen `setter`, der eine Temperatur in Celsius akzeptiert.

Denken Sie daran, dass `C = 5/9 * (F - 32)` und `F = C * 9.0 / 5 + 32` gelten, wobei `F` der Wert der Temperatur in Fahrenheit ist und `C` der Wert derselben Temperatur in Celsius.

**Hinweis:** Wenn Sie dies implementieren, verfolgen Sie die Temperatur innerhalb der Klasse in einer Skala, entweder Fahrenheit oder Celsius.

Das ist die Stärke eines Getters und Setters. Sie erstellen eine API für einen anderen Benutzer, der das korrekte Ergebnis erhält, unabhängig davon, welche Skala Sie verfolgen.

Mit anderen Worten abstrahieren Sie Implementierungsdetails vom Benutzer.

# --hints--

`Thermostat` sollte eine `class` mit einer definierten `constructor`-Methode sein.

```js
assert.isFunction(Thermostat);
assert.isFunction(Thermostat?.constructor);
```

Das Schlüsselwort `class` sollte verwendet werden.

```js
assert.match(code, /class/);
```

`Thermostat` sollte instanziierbar sein.

```js
const _t = new Thermostat(122);
assert.isObject(_t);
```

Wenn mit einem Fahrenheit-Wert instanziiert, sollte `Thermostat` die korrekte `temperature` setzen.

```js
const _t = new Thermostat(122);
assert.strictEqual(_t?.temperature, 50);
```

Ein `getter` sollte definiert sein.

```js
const _desc = Object.getOwnPropertyDescriptor(Thermostat.prototype, 'temperature');
assert.isFunction(_desc?.get);
```

Ein `setter` sollte definiert sein.

```js
const _desc = Object.getOwnPropertyDescriptor(Thermostat.prototype, 'temperature');
assert.isFunction(_desc?.set);
```

Das Aufrufen von `setter` mit einem Celsius-Wert sollte die `temperature` setzen.

```js
const _t = new Thermostat(32);
_t.temperature = 26;
const _u = new Thermostat(32);
_u.temperature = 50;
assert.approximately(_t.temperature, 26, 0.1);
assert.approximately(_u.temperature, 50, 0.1);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line

// Only change code above this line

const thermos = new Thermostat(76); // Setting in Fahrenheit scale
let temp = thermos.temperature; // 24.44 in Celsius
thermos.temperature = 26;
temp = thermos.temperature; // 26 in Celsius
```

# --solutions--

```js
class Thermostat {
  constructor(fahrenheit) {
    this._tempInCelsius = 5/9 * (fahrenheit - 32);
  }
  get temperature(){
    return this._tempInCelsius;
  }
  set temperature(newTemp){
    this._tempInCelsius = newTemp;
  }
}

const thermos = new Thermostat(76); // Setting in Fahrenheit scale
let temp = thermos.temperature; // 24.44 in Celsius
thermos.temperature = 26;
temp = thermos.temperature; // 26 in Celsius
```
