---
id: 587d7b87367417b2b2512b42
title: "Ein Array, das mit const deklariert wurde, verändern"
challengeType: 1
forumTopicId: 301206
dashedName: mutate-an-array-declared-with-const
---

# --description--

Wenn Sie mit `const` nicht vertraut sind, sehen Sie sich <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/declare-a-read-only-variable-with-the-const-keyword" target="_blank" rel="noopener noreferrer nofollow">diese Aufgabe zum <code>const</code>-Schlüsselwort</a> an.

Die `const`-Deklaration hat viele Anwendungsfälle im modernen JavaScript.

Einige Entwickler bevorzugen es, alle ihre Variablen standardmäßig mit `const` zuzuweisen, es sei denn, sie wissen, dass sie den Wert neu zuweisen müssen. Nur in diesem Fall verwenden sie `let`.

Es ist jedoch wichtig zu verstehen, dass Objekte (einschließlich Arrays und Funktionen), die einer Variablen mit `const` zugewiesen werden, weiterhin veränderbar sind. Die Verwendung der `const`-Deklaration verhindert nur die Neu-Zuweisung des Variablenbezeichners.

```js
const s = [5, 6, 7];
s = [1, 2, 3];
s[2] = 45;
console.log(s);
```

`s = [1, 2, 3]` führt zu einem Fehler. Nachdem Sie diese Zeile auskommentiert haben, zeigt `console.log` den Wert `[5, 6, 45]` an.

Wie Sie sehen, können Sie das Objekt `[5, 6, 7]` selbst verändern, und die Variable `s` zeigt weiterhin auf das veränderte Array `[5, 6, 45]`. Wie alle Arrays sind die Array-Elemente in `s` veränderbar, aber weil `const` verwendet wurde, können Sie den Variablenbezeichner `s` nicht verwenden, um mit dem Zuweisungsoperator auf ein anderes Array zu zeigen.

# --instructions--

Ein Array ist als `const s = [5, 7, 2]` deklariert. Ändern Sie das Array zu `[2, 5, 7]` durch verschiedene Elementzuweisungen.

# --hints--

Sie sollten das Schlüsselwort `const` nicht ersetzen.

```js
assert(__helpers.removeJSComments(code).match(/const/g));
```

`s` sollte eine konstante Variable sein (durch Verwendung von `const`).

```js
assert(__helpers.removeJSComments(code).match(/const\s+s/g));
```

Sie sollten die ursprüngliche Array-Deklaration nicht ändern.

```js
assert(__helpers.removeJSComments(code).match(
/const\s+s\s*=\s*\[\s*5\s*,\s*7\s*,\s*2\s*\]\s*;?/g
));
```

`s` sollte gleich `[2, 5, 7]` sein.

```js
assert.deepEqual(s, [2, 5, 7]);
```

# --seed--

## --seed-contents--

```js
const s = [5, 7, 2];
function editInPlace() {
  // Only change code below this line

  // Using s = [2, 5, 7] would be invalid

  // Only change code above this line
}
editInPlace();
```

# --solutions--

```js
const s = [5, 7, 2];
function editInPlace() {
  s[0] = 2;
  s[1] = 5;
  s[2] = 7;
}
editInPlace();
```
