---
id: 587d7b8c367417b2b2512b57
title: Verwenden Sie * zum Importieren aller Inhalte aus einer Datei
challengeType: 1
forumTopicId: 301210
dashedName: use--to-import-everything-from-a-file
---

# --description--

Angenommen, Sie haben eine Datei und möchten alle ihre Inhalte in die aktuelle Datei importieren. Dies kann mit der `import * as`-Syntax gemacht werden. Hier ist ein Beispiel, bei dem die Inhalte einer Datei namens `math_functions.js` in eine Datei im gleichen Verzeichnis importiert werden:

```js
import * as myMathModule from "./math_functions.js";
```

Die obige `import`-Anweisung erstellt ein Objekt namens `myMathModule`. Dies ist nur ein Variablenname, Sie können es beliebig benennen. Das Objekt enthält alle Exporte von `math_functions.js`, sodass Sie auf die Funktionen zugreifen können, wie Sie es bei jeder anderen Objekteigenschaft tun würden. So können Sie die Funktionen `add` und `subtract` verwenden, die importiert wurden:

```js
myMathModule.add(2,3);
myMathModule.subtract(5,3);
```

# --instructions--

Der Code in dieser Datei benötigt die Inhalte der Datei `string_functions.js`, die sich im gleichen Verzeichnis wie die aktuelle Datei befindet. Verwenden Sie die `import * as`-Syntax, um alles aus der Datei in ein Objekt namens `stringFunctions` zu importieren.

# --hints--

Ihr Code sollte die `import * as`-Syntax korrekt verwenden.

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*\*\s*as\s+stringFunctions\s+from\s*('|")\.\/string_functions\.js\1/g
  )
);
```

# --seed--

## --seed-contents--

```js

// Only change code above this line

stringFunctions.uppercaseString("hello");
stringFunctions.lowercaseString("WORLD!");
```

# --solutions--

```js
import * as stringFunctions from "./string_functions.js";

// add code above this line
stringFunctions.uppercaseString("hello");
stringFunctions.lowercaseString("WORLD!");
```
