---
id: 587d7b8a367417b2b2512b4c
title: Destrukturierung mit Rest-Elementen
challengeType: 1
forumTopicId: 301218
dashedName: use-destructuring-assignment-with-the-rest-parameter-to-reassign-array-elements
---

# --description--

In manchen Situationen bei der Destrukturierung von Arrays möchten wir vielleicht die restlichen Elemente in einem separaten Array sammeln.

Das Ergebnis ist ähnlich wie `Array.prototype.slice()`, wie unten gezeigt:

```js
const [a, b, ...arr] = [1, 2, 3, 4, 5, 7];
console.log(a, b);
console.log(arr);
```

Die Konsole würde die Werte `1, 2` und `[3, 4, 5, 7]` anzeigen.

Die Variablen `a` und `b` nehmen die ersten beiden Werte aus dem Array. Danach erhält `arr` aufgrund der Rest-Syntax den Rest der Werte in Form eines Arrays. Das Rest-Element funktioniert nur korrekt als letzte Variable in der Liste. Das heißt, Sie können die Rest-Syntax nicht verwenden, um ein Teilarray zu erfassen, das das letzte Element des ursprünglichen Arrays auslässt.

# --instructions--

Verwenden Sie eine Destrukturierungszuweisung mit der Rest-Syntax, um das Verhalten von `Array.prototype.slice()` zu emulieren. `removeFirstTwo()` soll ein Teilarray des ursprünglichen Arrays `list` zurückgeben, bei dem die ersten beiden Elemente ausgelassen sind.

# --hints--

`removeFirstTwo([1, 2, 3, 4, 5])` sollte `[3, 4, 5]` sein

```js
assert.deepEqual(removeFirstTwo([1, 2, 3, 4, 5]), [3, 4, 5]);
```

`removeFirstTwo()` sollte `list` nicht verändern

```js
const _testArr = [1, 2, 3, 4, 5];
removeFirstTwo(_testArr);
assert.deepEqual(_testArr, [1, 2, 3, 4, 5])
```

`Array.slice()` sollte nicht verwendet werden.

```js
assert(!__helpers.removeJSComments(code).match(/\.\s*slice\s*\(/));
```

Sie sollten die Rest-Syntax verwenden.

```js
assert.match(code, /\.\.\./);
```

# --seed--

## --seed-contents--

```js
function removeFirstTwo(list) {
  return list;
}

const source = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const sourceWithoutFirstTwo = removeFirstTwo(source);
```

# --solutions--

```js
function removeFirstTwo(list) {
  // comment with 'slice' to check comments are removed in tests
  const [, , ...shorterList] = list;
  return shorterList;
}

const source = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const sourceWithoutFirstTwo = removeFirstTwo(source);
```
