---
id: 587d7b8b367417b2b2512b50
title: Schreibe prägnante deklarative Funktionen mit ES6
challengeType: 1
forumTopicId: 301224
dashedName: write-concise-declarative-functions-with-es6
---

# --description--

Beim Definieren von Funktionen innerhalb von Objekten in ES5 müssen wir das Schlüsselwort `function` wie folgt verwenden:

```js
const person = {
  name: "Taylor",
  sayHello: function() {
    return `Hello! My name is ${this.name}.`;
  }
};
```

Mit ES6 können Sie das Schlüsselwort `function` und den Doppelpunkt beim Definieren von Funktionen in Objekten ganz weglassen. Hier ist ein Beispiel für diese Syntax:

```js
const person = {
  name: "Taylor",
  sayHello() {
    return `Hello! My name is ${this.name}.`;
  }
};
```

# --instructions--

Überarbeiten Sie die Funktion `setGear` innerhalb des Objekts `bicycle`, sodass die oben beschriebene Kurzschreibweise verwendet wird.

# --hints--

Die traditionelle Funktionsausdruck-Syntax sollte nicht verwendet werden.

```js
assert(!__helpers.removeJSComments(code).match(/function/));
```

`setGear` sollte eine deklarative Funktion sein.

```js
assert(
  typeof bicycle.setGear === 'function' && __helpers.removeJSComments(code).match(/setGear\s*\(.+\)\s*\{/)
);
```

`bicycle.setGear(48)` sollte den Wert von `gear` auf 48 ändern.

```js
bicycle.setGear(48);
assert(bicycle.gear === 48);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const bicycle = {
  gear: 2,
  setGear: function(newGear) {
    this.gear = newGear;
  }
};
// Only change code above this line
bicycle.setGear(3);
console.log(bicycle.gear);
```

# --solutions--

```js
const bicycle = {
  gear: 2,
  // setGear: function(newGear) {
  setGear(newGear) {
    this.gear = newGear;
  }
};
bicycle.setGear(3);
```
