---
id: 587d7b89367417b2b2512b49
title: "Verwenden Sie Destrukturierungszuweisung, um Variablen aus Objekten zuzuweisen"
challengeType: 1
forumTopicId: 301215
dashedName: use-destructuring-assignment-to-assign-variables-from-objects
---

# --description--

Destrukturierung ermöglicht es Ihnen, beim Extrahieren von Werten einen neuen Variablennamen zuzuweisen. Sie können dies tun, indem Sie den neuen Namen nach einem Doppelpunkt angeben, wenn Sie den Wert zuweisen.

Verwenden Sie dasselbe Objekt aus dem letzten Beispiel:

```js
const user = { name: 'John Doe', age: 34 };
```

So können Sie in der Zuweisung neue Variablennamen vergeben:

```js
const { name: userName, age: userAge } = user;
```

Sie können es so lesen: „Hole den Wert von `user.name` und weise ihn einer neuen Variablen namens `userName` zu“ und so weiter. Der Wert von `userName` wäre der String `John Doe`, und der Wert von `userAge` wäre die Zahl `34`.

# --instructions--

Ersetzen Sie die beiden Zuweisungen durch eine äquivalente Destrukturierungszuweisung. Sie soll weiterhin den Variablen `highToday` und `highTomorrow` die Werte von `today` und `tomorrow` aus dem `HIGH_TEMPERATURES`-Objekt zuweisen.

# --hints--

Sie sollten die ES5-Zuweisungssyntax entfernen.

```js
assert(
  !__helpers.removeJSComments(code).match(/highToday = HIGH_TEMPERATURES\.today/g) &&
    !__helpers.removeJSComments(code).match(/highTomorrow = HIGH_TEMPERATURES\.tomorrow/g)
);
```

Sie sollten Destrukturierung verwenden, um die Variable `highToday` zu erstellen.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*(today\s*:\s*highToday[^}]*|[^,]*,\s*today\s*:\s*highToday\s*)}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g
  )
);
```

Sie sollten Destrukturierung verwenden, um die Variable `highTomorrow` zu erstellen.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*(tomorrow\s*:\s*highTomorrow[^}]*|[^,]*,\s*tomorrow\s*:\s*highTomorrow\s*)}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g
  )
);
```

`highToday` sollte gleich `77` sein und `highTomorrow` sollte gleich `80` sein.

```js
assert(highToday === 77 && highTomorrow === 80);
```

# --seed--

## --seed-contents--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

// Only change code below this line
  
const highToday = HIGH_TEMPERATURES.today;
const highTomorrow = HIGH_TEMPERATURES.tomorrow; 

// Only change code above this line
```

# --solutions--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

const { today: highToday, tomorrow: highTomorrow } = HIGH_TEMPERATURES;
```
