---
id: 587d7b8a367417b2b2512b4e
title: Strings mit Template-Literalen erstellen
challengeType: 1
forumTopicId: 301200
dashedName: create-strings-using-template-literals
---

# --description--

Eine neue Funktion von ES6 ist das <dfn>Template-Literal</dfn>. Dabei handelt es sich um eine spezielle Art von String, die das Erstellen komplexer Strings erleichtert.

Template-Literale ermöglichen es Ihnen, mehrzeilige Strings zu erstellen und String-Interpolation zu verwenden, um Strings zu erzeugen.

Betrachten Sie den folgenden Code:

```js
const person = {
  name: "Zodiac Hasbro",
  age: 56
};

const greeting = `Hello, my name is ${person.name}!
I am ${person.age} years old.`;

console.log(greeting);
```

Die Konsole zeigt die Strings `Hello, my name is Zodiac Hasbro!` und `I am 56 years old.` an.

Dort ist viel passiert. Zuerst verwendet das Beispiel Backticks (`` ` ``), not quotes (`'` or `"`), to wrap the string. Secondly, notice that the string is multi-line, both in the code and the output. This saves inserting `\n` within strings. The `${variable}` syntax used above is a placeholder. Basically, you won't have to use concatenation with the `+` operator anymore. To add variables to strings, you just drop the variable in a template string and wrap it with `${` and `}`. Similarly, you can include other expressions in your string literal, for example `${a + b}`). Diese neue Art, Strings zu erstellen, gibt Ihnen mehr Flexibilität, um robuste Strings zu erzeugen.

# --instructions--

Verwenden Sie die Syntax der Template-Literale mit Backticks, um ein Array von Listenelement-Strings (`li`) zu erstellen. Der Text jedes Listenelements soll eines der Array-Elemente aus der `failure`-Eigenschaft des `result`-Objekts sein und ein `class`-Attribut mit dem Wert `text-warning` besitzen. Die `makeList`-Funktion soll das Array der Listenelement-Strings zurückgeben.

Verwenden Sie eine Iterator-Methode (eine beliebige Schleife), um die gewünschte Ausgabe (unten gezeigt) zu erhalten.

```js
[
  '<li class="text-warning">no-var</li>',
  '<li class="text-warning">var-on-top</li>',
  '<li class="text-warning">linebreak</li>'
]
```

# --hints--

`failuresList` sollte ein Array mit `result failure`-Nachrichten sein.

```js
assert(
  typeof makeList(result.failure) === 'object' && failuresList.length === 3
);
```

`failuresList` sollte der angegebenen Ausgabe entsprechen.

```js
assert(
  makeList(result.failure).every(
    (v, i) =>
      v === `<li class="text-warning">${result.failure[i]}</li>` ||
      v === `<li class='text-warning'>${result.failure[i]}</li>`
  )
);
```

Template-Strings und Ausdrucksinterpolation sollten verwendet werden.

```js
assert.match(code, /(`.*\${.*}.*`)/);
```

Ein Iterator sollte verwendet werden.

```js
assert(__helpers.removeJSComments(code).match(/for|map|reduce|forEach|while/));
```

# --seed--

## --seed-contents--

```js
const result = {
  success: ["max-length", "no-amd", "prefer-arrow-functions"],
  failure: ["no-var", "var-on-top", "linebreak"],
  skipped: ["no-extra-semi", "no-dup-keys"]
};
function makeList(arr) {
  // Only change code below this line
  const failureItems = [];
  // Only change code above this line

  return failureItems;
}

const failuresList = makeList(result.failure);
```

# --solutions--

```js
const result = {
  success: ["max-length", "no-amd", "prefer-arrow-functions"],
  failure: ["no-var", "var-on-top", "linebreak"],
  skipped: ["no-extra-semi", "no-dup-keys"]
};
function makeList(arr) {
  return arr.map(val => `<li class="text-warning">${val}</li>`);
}

const failuresList = makeList(result.failure);
```
