---
id: 587d7b89367417b2b2512b4a
title: "Verwenden Sie Destrukturierungszuweisung, um Variablen aus verschachtelten Objekten zuzuweisen"
challengeType: 1
forumTopicId: 301214
dashedName: use-destructuring-assignment-to-assign-variables-from-nested-objects
---

# --description--

Sie können die gleichen Grundsätze aus den vorherigen beiden Lektionen verwenden, um Werte aus verschachtelten Objekten zu destrukturieren.

Wenn Sie ein Objekt ähnlich den vorherigen Beispielen verwenden:

```js
const user = {
  johnDoe: { 
    age: 34,
    email: 'johnDoe@freeCodeCamp.com'
  }
};
```

So können Sie die Werte von Objekteigenschaften extrahieren und Variablen mit demselben Namen zuweisen:

```js
const { johnDoe: { age, email }} = user;
```

Und so können Sie die Werte von Objekteigenschaften Variablen mit anderen Namen zuweisen:

```js
const { johnDoe: { age: userAge, email: userEmail }} = user;
```

# --instructions--

Ersetzen Sie die beiden Zuweisungen durch eine äquivalente Destrukturierungszuweisung. Sie soll weiterhin den Variablen `lowToday` und `highToday` die Werte von `today.low` und `today.high` aus dem `LOCAL_FORECAST` Objekt zuweisen.

# --hints--

Sie sollten die ES5-Zuweisungssyntax entfernen.

```js
assert(
  !__helpers.removeJSComments(code).match(/lowToday = LOCAL_FORECAST\.today\.low/g) &&
    !__helpers.removeJSComments(code).match(/highToday = LOCAL_FORECAST\.today.high/g)
);
```

Sie sollten Destrukturierung verwenden, um die Variable `lowToday` zu erstellen.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*today\s*:\s*{\s*(low\s*:\s*lowToday[^}]*|[^,]*,\s*low\s*:\s*lowToday\s*)},?\s*}\s*=\s*LOCAL_FORECAST(;|\s+|\/\/)/g
  )
);
```

Sie sollten Destrukturierung verwenden, um die Variable `highToday` zu erstellen.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*today\s*:\s*{\s*(high\s*:\s*highToday[^}]*|[^,]*,\s*high\s*:\s*highToday,?\s*)},?\s*}\s*=\s*LOCAL_FORECAST(;|\s+|\/\/)/g
  )
);
```

`lowToday` sollte gleich `64` sein und `highToday` sollte gleich `77` sein.

```js
assert(lowToday === 64 && highToday === 77);
```

# --seed--

## --seed-contents--

```js
const LOCAL_FORECAST = {
  yesterday: { low: 61, high: 75 },
  today: { low: 64, high: 77 },
  tomorrow: { low: 68, high: 80 }
};

// Only change code below this line
  
const lowToday = LOCAL_FORECAST.today.low;
const highToday = LOCAL_FORECAST.today.high;

// Only change code above this line
```

# --solutions--

```js
const LOCAL_FORECAST = {
  yesterday: { low: 61, high: 75 },
  today: { low: 64, high: 77 },
  tomorrow: { low: 68, high: 80 }
};
 
const { today: { low: lowToday, high: highToday }} = LOCAL_FORECAST;
```
