---
id: 587d7b8d367417b2b2512b59
title: Einen Standardexport importieren
challengeType: 1
forumTopicId: 301205
dashedName: import-a-default-export
---

# --description--

In der letzten Aufgabe haben Sie etwas über `export default` und dessen Verwendung gelernt. Um einen Standardexport zu importieren, müssen Sie eine andere `import`-Syntax verwenden. Im folgenden Beispiel ist `add` der Standardexport der `math_functions.js`-Datei. So importieren Sie ihn:

```js
import add from "./math_functions.js";
```

Die Syntax unterscheidet sich an einer entscheidenden Stelle. Der importierte Wert, `add`, ist nicht von geschweiften Klammern (`{}`) umgeben. `add` ist hier einfach ein Variablenname für den Standardexport der `math_functions.js`-Datei. Sie können beim Import eines Standardexports jeden beliebigen Namen verwenden.

# --instructions--

Importieren Sie im folgenden Code den Standardexport aus der `math_functions.js`-Datei, die sich im selben Verzeichnis wie diese Datei befindet. Geben Sie dem Import den Namen `subtract`.

# --hints--

Sie sollten `subtract` korrekt aus `math_functions.js` importieren.

```js
assert(__helpers.removeJSComments(code).match(/import\s+subtract\s+from\s+('|")\.\/math_functions\.js\1/g));
```

# --seed--

## --seed-contents--

```js
  
// Only change code above this line

subtract(7,4);
```

# --solutions--

```js
import subtract from "./math_functions.js";

subtract(7,4);
```
