---
id: 587d7b89367417b2b2512b4b
title: "Verwenden Sie Destrukturierungszuweisung, um Variablen aus Arrays zuzuweisen"
challengeType: 1
forumTopicId: 301213
dashedName: use-destructuring-assignment-to-assign-variables-from-arrays
---

# --description--

ES6 macht die Destrukturierung von Arrays genauso einfach wie die Destrukturierung von Objekten.

Ein wesentlicher Unterschied zwischen dem Spread-Operator und der Array-Destrukturierung besteht darin, dass der Spread-Operator alle Inhalte eines Arrays in eine durch Kommas getrennte Liste entpackt. Folglich können Sie nicht auswählen, welche Elemente Sie Variablen zuweisen möchten.

Die Destrukturierung eines Arrays ermöglicht genau das:

```js
const [a, b] = [1, 2, 3, 4, 5, 6];
console.log(a, b);
```

Die Konsole zeigt die Werte von `a` und `b` als `1, 2` an.

Der Variablen `a` wird der erste Wert des Arrays zugewiesen, und `b` wird der zweite Wert des Arrays zugewiesen. Wir können auch auf den Wert an einem beliebigen Index in einem Array mit Destrukturierung zugreifen, indem wir Kommas verwenden, um den gewünschten Index zu erreichen:

```js
const [a, b,,, c] = [1, 2, 3, 4, 5, 6];
console.log(a, b, c);
```

Die Konsole zeigt die Werte von `a`, `b` und `c` als `1, 2, 5` an.

# --instructions--

Verwenden Sie die Destrukturierungszuweisung, um die Werte von `a` und `b` zu tauschen, sodass `a` den Wert erhält, der in `b` gespeichert ist, und `b` den Wert, der in `a` gespeichert ist.

# --hints--

Der Wert von `a` sollte nach dem Tausch `6` sein.

```js
assert(a === 6);
```

Der Wert von `b` sollte nach dem Tausch `8` sein.

```js
assert(b === 8);
```

Sie sollten Array-Destrukturierung verwenden, um `a` und `b` zu tauschen.

```js
assert(/\[\s*(\w)\s*,\s*(\w)\s*\]\s*=\s*\[\s*\2\s*,\s*\1\s*\]/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let a = 8, b = 6;
// Only change code below this line
```

# --solutions--

```js
let a = 8, b = 6;
[a, b] = [b, a];
```
