---
id: 587d7b8a367417b2b2512b4f
title: Schreibe prägnante Objektliteral-Deklarationen mit Objekt-Eigenschafts-Shorthand
challengeType: 1
forumTopicId: 301225
dashedName: write-concise-object-literal-declarations-using-object-property-shorthand
---

# --description--

ES6 bietet eine praktische Unterstützung, um Objektliterale einfach zu definieren.

Betrachten Sie den folgenden Code:

```js
const getMousePosition = (x, y) => ({
  x: x,
  y: y
});
```

`getMousePosition` ist eine einfache Funktion, die ein Objekt mit zwei Eigenschaften zurückgibt. ES6 stellt syntaktischen Zucker bereit, um die Redundanz zu vermeiden, `x: x` zweimal schreiben zu müssen. Sie können einfach `x` einmal schreiben, und es wird unter der Haube in `x: x` (oder etwas Ähnliches) umgewandelt. Hier ist dieselbe Funktion von oben, umgeschrieben mit dieser neuen Syntax:

```js
const getMousePosition = (x, y) => ({ x, y });
```

# --instructions--

Verwenden Sie Objekt-Eigenschafts-Shorthand bei Objektliteralen, um ein Objekt mit den Eigenschaften `name`, `age` und `gender` zu erstellen und zurückzugeben.

# --hints--

`createPerson("Zodiac Hasbro", 56, "male")` sollte `{name: "Zodiac Hasbro", age: 56, gender: "male"}` zurückgeben.

```js
assert.deepEqual(
  { name: 'Zodiac Hasbro', age: 56, gender: 'male' },
  createPerson('Zodiac Hasbro', 56, 'male')
);
```

Ihr Code sollte `key:value` nicht verwenden.

```js
assert(!__helpers.removeJSComments(code).match(/:/g))
```

# --seed--

## --seed-contents--

```js
const createPerson = (name, age, gender) => {
  // Only change code below this line
  return {
    name: name,
    age: age,
    gender: gender
  };
  // Only change code above this line
};
```

# --solutions--

```js
const createPerson = (name, age, gender) => {
  // Only change code below this line
  /*return {
    name: name,
    age: age,
    gender: gender
  };*/
  return {
    name,
    age,
    gender
  };
  // Only change code above this line
};
```
