---
id: 587d7b8b367417b2b2512b53
title: "Klassensyntax verwenden, um eine Konstruktorfunktion zu definieren"
challengeType: 1
forumTopicId: 301212
dashedName: use-class-syntax-to-define-a-constructor-function
---

# --description--

ES6 bietet eine neue Syntax zur Erstellung von Objekten mit dem SchlÃžsselwort <dfn>class</dfn>.

In ES5 kann ein Objekt erstellt werden, indem eine `constructor` Funktion definiert und das SchlÃžsselwort `new` verwendet wird, um das Objekt zu instanziieren.

In ES6 hat eine `class` Deklaration eine `constructor` Methode, die mit dem SchlÃžsselwort `new` aufgerufen wird. Wenn die `constructor` Methode nicht explizit definiert ist, wird sie implizit ohne Argumente definiert.

```js
// Explicit constructor
class SpaceShuttle {
  constructor(targetPlanet) {
    this.targetPlanet = targetPlanet;
  }
  takeOff() {
    console.log("To " + this.targetPlanet + "!");
  }
}

// Implicit constructor 
class Rocket {
  launch() {
    console.log("To the moon!");
  }
}

const zeus = new SpaceShuttle('Jupiter');
// prints To Jupiter! in console
zeus.takeOff();

const atlas = new Rocket();
// prints To the moon! in console
atlas.launch();
```

Es sollte beachtet werden, dass das SchlÃžsselwort `class` eine neue Funktion deklariert, zu der ein Konstruktor hinzugefÃžgt wird. Dieser Konstruktor wird aufgerufen, wenn `new` verwendet wird, um ein neues Objekt zu erstellen.

**Hinweis:** UpperCamelCase sollte konventionell fÃžr ES6-Klassennamen verwendet werden, wie oben bei `SpaceShuttle` gezeigt.

Die `constructor` Methode ist eine spezielle Methode zum Erstellen und Initialisieren eines mit einer Klasse erstellten Objekts. Mehr darÃžber lernen Sie im Abschnitt objektorientierte Programmierung der JavaScript-Algorithmen und Datenstrukturen Zertifizierung.

# --instructions--

Verwenden Sie das SchlÃžsselwort `class` und schreiben Sie eine `constructor`, um die `Vegetable` Klasse zu erstellen.

Die `Vegetable` Klasse ermÃķglicht es Ihnen, ein GemÃžseobjekt mit einer Eigenschaft `name` zu erstellen, die an den `constructor` Ãžbergeben wird.

# --hints--

`Vegetable` sollte eine `class` mit einer definierten `constructor` Methode sein.

```js
assert(
  typeof Vegetable === 'function' && typeof Vegetable.constructor === 'function'
);
```

Das SchlÃžsselwort `class` sollte verwendet werden.

```js
assert(__helpers.removeJSComments(code).match(/class/g));
```

`Vegetable` sollte instanziierbar sein.

```js
assert(() => {
  const a = new Vegetable('apple');
  return typeof a === 'object';
});
```

`carrot.name` sollte `carrot` zurÃžckgeben.

```js
assert(carrot.name == 'carrot');
```

# --seed--

## --seed-contents--

```js
// Only change code below this line

// Only change code above this line

const carrot = new Vegetable('carrot');
console.log(carrot.name); // Should display 'carrot'
```

# --solutions--

```js
class Vegetable {
  constructor(name) {
    this.name = name;
  }
}
const carrot = new Vegetable('carrot');
```
