---
id: 587d7b88367417b2b2512b46
title: Standardparameter für Ihre Funktionen festlegen
challengeType: 1
forumTopicId: 301209
dashedName: set-default-parameters-for-your-functions
---

# --description--

Um uns dabei zu helfen, flexiblere Funktionen zu erstellen, führt ES6 <dfn>Standardparameter</dfn> für Funktionen ein.

Sehen Sie sich diesen Code an:

```js
const greeting = (name = "Anonymous") => "Hello " + name;

console.log(greeting("John"));
console.log(greeting());
```

Die Konsole zeigt die Strings `Hello John` und `Hello Anonymous` an.

Der Standardparameter greift, wenn das Argument nicht angegeben ist (es ist undefined). Wie Sie im obigen Beispiel sehen, erhält der Parameter `name` seinen Standardwert `Anonymous`, wenn Sie keinen Wert für den Parameter angeben. Sie können für beliebig viele Parameter Standardwerte hinzufügen.

# --instructions--

Ändern Sie die Funktion `increment`, indem Sie Standardparameter hinzufügen, sodass 1 zu `number` addiert wird, falls `value` nicht angegeben ist.

# --hints--

Das Ergebnis von `increment(5, 2)` sollte `7` sein.

```js
assert(increment(5, 2) === 7);
```

Das Ergebnis von `increment(5)` sollte `6` sein.

```js
assert(increment(5) === 6);
```

Ein Standardparameterwert von `1` sollte für `value` verwendet werden.

```js
assert(__helpers.removeJSComments(code).match(/value\s*=\s*1/g));
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const increment = (number, value) => number + value;
// Only change code above this line
```

# --solutions--

```js
const increment = (number, value = 1) => number + value;
```
