---
id: 587d7b8c367417b2b2512b58
title: Einen Export-Fallback mit export default erstellen
challengeType: 1
forumTopicId: 301199
dashedName: create-an-export-fallback-with-export-default
---

# --description--

In der `export`-Lektion haben Sie die Syntax kennengelernt, die als <dfn>named export</dfn> bezeichnet wird. Damit konnten Sie mehrere Funktionen und Variablen für die Verwendung in anderen Dateien verfügbar machen.

Es gibt eine weitere `export`-Syntax, die Sie kennen müssen, bekannt als <dfn>export default</dfn>. Diese Syntax verwenden Sie normalerweise, wenn nur ein Wert aus einer Datei exportiert wird. Sie wird auch verwendet, um einen Fallback-Wert für eine Datei oder ein Modul zu erstellen.

Nachfolgend sehen Sie Beispiele mit `export default`:

```js
export default function add(x, y) {
  return x + y;
}

export default function(x, y) {
  return x + y;
}
```

Das erste ist eine benannte Funktion, und das zweite ist eine anonyme Funktion.

Da `export default` verwendet wird, um einen Fallback-Wert für ein Modul oder eine Datei zu deklarieren, kann in jedem Modul oder jeder Datei nur ein Wert als Default-Export vorhanden sein. Außerdem können Sie `export default` nicht zusammen mit `var`, `let` oder `const` verwenden.

# --instructions--

Die folgende Funktion soll der Fallback-Wert für das Modul sein. Bitte fügen Sie den notwendigen Code hinzu, um dies zu erreichen.

# --hints--

Ihr Code sollte einen `export`-Fallback verwenden.

```js
assert(
  __helpers.removeJSComments(code).match(
    /export\s+default\s+function(\s+subtract\s*|\s*)\(\s*x,\s*y\s*\)\s*{/g
  )
);
```

# --seed--

## --seed-contents--

```js
function subtract(x, y) {
  return x - y;
}
```

# --solutions--

```js
export default function subtract(x, y) {
  return x - y;
}
```
