---
id: 598f48a36c8c40764b4e52b3
title: Mutation von Objekten verhindern
challengeType: 1
forumTopicId: 301207
dashedName: prevent-object-mutation
---

# --description--

Wie in der vorherigen Aufgabe zu sehen war, schützt die Deklaration von `const` allein Ihre Daten nicht wirklich vor Mutation. Um sicherzustellen, dass Ihre Daten sich nicht ändern, stellt JavaScript die Funktion `Object.freeze` bereit, um Datenmutation zu verhindern.

Jeder Versuch, das Objekt zu ändern, wird abgelehnt, wobei ein Fehler geworfen wird, wenn das Skript im Strict-Modus ausgeführt wird.

```js
let obj = {
  name:"FreeCodeCamp",
  review:"Awesome"
};
Object.freeze(obj);
obj.review = "bad";
obj.newProp = "Test";
console.log(obj); 
```

Die Zuweisungen `obj.review` und `obj.newProp` führen zu Fehlern, weil unser Editor standardmäßig im Strict-Modus läuft, und die Konsole zeigt den Wert `{ name: "FreeCodeCamp", review: "Awesome" }` an.

# --instructions--

In dieser Aufgabe verwenden Sie `Object.freeze`, um zu verhindern, dass sich mathematische Konstanten ändern. Sie müssen das Objekt `MATH_CONSTANTS` einfrieren, damit niemand den Wert von `PI` ändern, Eigenschaften hinzufügen oder löschen kann.

# --hints--

Sie sollten das Schlüsselwort `const` nicht ersetzen.

```js
assert(__helpers.removeJSComments(code).match(/const/g));
```

`MATH_CONSTANTS` sollte eine konstante Variable sein (durch Verwendung von `const`).

```js
assert(__helpers.removeJSComments(code).match(/const\s+MATH_CONSTANTS/g));
```

Sie sollten die ursprüngliche Deklaration von `MATH_CONSTANTS` nicht ändern.

```js
assert(__helpers.removeJSComments(code).match(
   /const\s+MATH_CONSTANTS\s+=\s+{\s+PI:\s+3.14\s+};/g
));
```

`PI` sollte gleich `3.14` sein.

```js
assert(PI === 3.14);
```

# --seed--

## --seed-contents--

```js
function freezeObj() {
  const MATH_CONSTANTS = {
    PI: 3.14
  };
  // Only change code below this line


  // Only change code above this line
  try {
    MATH_CONSTANTS.PI = 99;
  } catch(ex) {
    console.log(ex);
  }
  return MATH_CONSTANTS.PI;
}
const PI = freezeObj();
```

# --solutions--

```js
function freezeObj() {
  const MATH_CONSTANTS = {
    PI: 3.14
  };
  Object.freeze(MATH_CONSTANTS);

  try {
    MATH_CONSTANTS.PI = 99;
  } catch(ex) {
    console.log(ex);
  }
  return MATH_CONSTANTS.PI;
}
const PI = freezeObj();
```
