---
id: 5cdafbb0291309899753167f
title: Erstellen Sie ein JavaScript-Promise
challengeType: 1
forumTopicId: 301197
dashedName: create-a-javascript-promise
---

# --description--

Ein Promise in JavaScript ist genau das, was der Name sagt – Sie verwenden es, um ein Versprechen zu geben, etwas zu tun, normalerweise asynchron. Wenn die Aufgabe abgeschlossen ist, erfüllen Sie entweder Ihr Versprechen oder nicht. `Promise` ist eine Konstruktorfunktion, daher müssen Sie das Schlüsselwort `new` verwenden, um eine Instanz zu erstellen. Es nimmt eine Funktion als Argument, mit zwei Parametern – `resolve` und `reject`. Diese sind Methoden, die verwendet werden, um das Ergebnis des Promise zu bestimmen. Die Syntax sieht so aus:

```js
const myPromise = new Promise((resolve, reject) => {

});
```

# --instructions--

Erstellen Sie ein neues Promise mit dem Namen `makeServerRequest`. Übergeben Sie dem Konstruktor eine Funktion mit den Parametern `resolve` und `reject`.

# --hints--

Sie sollten ein Promise einer deklarierten Variablen namens `makeServerRequest` zuweisen.

```js
assert(makeServerRequest instanceof Promise);
```

Ihr Promise sollte eine Funktion mit den Parametern `resolve` und `reject` erhalten.

```js
assert(
  __helpers.removeJSComments(code).match(
    /Promise\s*\(\s*(function\s*\(\s*resolve\s*,\s*reject\s*\)\s*{|\(\s*resolve\s*,\s*reject\s*\)\s*=>\s*{)[^}]*}/g
  )
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {

});
```
