---
id: 5cdafbc32913098997531680
title: Promise mit resolve und reject abschließen
challengeType: 1
forumTopicId: 301196
dashedName: complete-a-promise-with-resolve-and-reject
---

# --description--

Ein Promise hat drei Zustände: `pending`, `fulfilled` und `rejected`. Das Promise, das Sie in der letzten Aufgabe erstellt haben, steckt für immer im Zustand `pending`, weil Sie keine Möglichkeit hinzugefügt haben, das Promise abzuschließen. Die Parameter `resolve` und `reject`, die dem Promise-Argument übergeben werden, dienen genau dazu. `resolve` wird verwendet, wenn Ihr Promise erfolgreich sein soll, und `reject`, wenn es fehlschlagen soll. Dabei handelt es sich um Methoden, die ein Argument entgegennehmen, wie unten gezeigt.

```js
const myPromise = new Promise((resolve, reject) => {
  if(condition here) {
    resolve("Promise was fulfilled");
  } else {
    reject("Promise was rejected");
  }
});
```

Im obigen Beispiel werden Strings als Argumente für diese Funktionen verwendet, aber es kann tatsächlich alles Mögliche sein. Oft ist es ein Objekt, aus dem Sie Daten entnehmen, um sie auf Ihrer Website oder anderswo zu verwenden.

# --instructions--

Lassen Sie das Promise Erfolg und Fehler behandeln. Wenn `responseFromServer` `true` ist, rufen Sie die Methode `resolve` auf, um das Promise erfolgreich abzuschließen. Übergeben Sie `resolve` einen String mit dem Wert `We got the data`. Wenn `responseFromServer` `false` ist, verwenden Sie stattdessen die Methode `reject` und übergeben ihr den String: `Data not received`.

# --hints--

`resolve` sollte mit dem erwarteten String aufgerufen werden, wenn die Bedingung `if` `true` ist.

```js
assert(
  __helpers.removeJSComments(code).match(/if\s*\(\s*responseFromServer\s*\)\s*{\s*resolve\s*\(\s*('|"|`)We got the data\1\s*\)(\s*|\s*;\s*)}/g)
);
```

`reject` sollte mit dem erwarteten String aufgerufen werden, wenn die Bedingung `if` `false` ist.

```js
assert(
  __helpers.removeJSComments(code).match(/}\s*else\s*{\s*reject\s*\(\s*('|"|`)Data not received\1\s*\)(\s*|\s*;\s*)}/g)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer represents a response from a server
  let responseFromServer;
    
  if(responseFromServer) {
    // Change this line
  } else {  
    // Change this line
  }
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer represents a response from a server
  let responseFromServer;

  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});
```
