---
id: 5cdafbe72913098997531682
title: Mit catch eine abgelehnte Promise behandeln
challengeType: 1
forumTopicId: 301204
dashedName: handle-a-rejected-promise-with-catch
---

# --description--

`catch` ist die Methode, die verwendet wird, wenn Ihre Promise abgelehnt wurde. Sie wird unmittelbar ausgeführt, nachdem die `reject`-Methode einer Promise aufgerufen wurde. Hier ist die Syntax:

```js
myPromise.catch(error => {
  
});
```

`error` ist das Argument, das an die `reject`-Methode übergeben wird.

# --instructions--

Fügen Sie Ihrer Promise die `catch`-Methode hinzu. Verwenden Sie `error` als Parameter der Callback-Funktion und geben Sie `error` in der Konsole aus.

# --hints--

Sie sollten die `catch`-Methode für die Promise aufrufen.

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/(makeServerRequest|\))\.catch\(/g)
);
```

Ihre `catch`-Methode sollte eine Callback-Funktion mit `error` als Parameter haben.

```js
assert(
  /\.catch\((function\(error\){|error|\(error\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

Sie sollten `error` in der Konsole ausgeben.

```js
assert(
  /\.catch\((function\(error\){|error|\(error\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  ) &&
    __helpers
      .removeWhiteSpace(__helpers.removeJSComments(code))
      .match(/\.catch\(.*?error.*?console.log\(error\).*?\)/)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to false to represent an unsuccessful response from a server
  let responseFromServer = false;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to false to represent an unsuccessful response from a server
  let responseFromServer = false;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});

makeServerRequest.catch(error => {
  console.log(error);
});
```
