---
id: 5cdafbd72913098997531681
title: Mit then eine erfüllte Promise verarbeiten
challengeType: 1
forumTopicId: 301203
dashedName: handle-a-fulfilled-promise-with-then
---

# --description--

Promises sind besonders nützlich, wenn Sie in Ihrem Code einen Prozess haben, der eine unbekannte Zeit in Anspruch nimmt (also etwas Asynchrones), oft eine Anfrage an einen Server. Wenn Sie eine Anfrage an einen Server stellen, dauert das eine gewisse Zeit, und nachdem sie abgeschlossen ist, möchten Sie normalerweise etwas mit der Antwort des Servers machen. Dies kann durch die Verwendung der `then`-Methode erreicht werden.

```js
Promise.prototype.then(onFulfilled, onRejected)
```

Die `then`-Methode plant Callback-Funktionen für den endgültigen Abschluss einer Promise – entweder Erfüllung oder Ablehnung. Einer der `onFulfilled`- und `onRejected`-Handler wird ausgeführt, um die Erfüllung oder Ablehnung der aktuellen Promise zu verarbeiten. Wenn die Promise mit `resolve` erfüllt wird, wird der `onFulfilled`-Handler aufgerufen.

```js
myPromise.then(result => {
  
});
```

`result` stammt aus dem Argument, das der `resolve`-Methode übergeben wird.

# --instructions--

Fügen Sie die `then`-Methode zu Ihrer Promise hinzu. Verwenden Sie `result` als Parameter der Callback-Funktion und geben Sie `result` in der Konsole aus.

# --hints--

Sie sollten die `then`-Methode auf der Promise aufrufen.

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/(makeServerRequest|\))\.then\(/g)
);
```

Ihre `then`-Methode sollte eine Callback-Funktion mit `result` als Parameter haben.

```js
assert(
  /\.then\((function\(result\){|result|\(result\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

Sie sollten `result` in der Konsole ausgeben.

```js
assert(
  /\.then\((function\(result\){|result|\(result\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  ) &&
    __helpers
      .removeWhiteSpace(__helpers.removeJSComments(code))
      .match(/\.then\(.*?result.*?console.log\(result\).*?\)/)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to true to represent a successful response from a server
  let responseFromServer = true;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to true to represent a successful response from a server
  let responseFromServer = true;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});
```
