---
id: 5cddbfd622f1a59093ec611d
title: Ein Modulskript erstellen
challengeType: 6
forumTopicId: 301198
dashedName: create-a-module-script
---

# --description--

JavaScript begann mit einer kleinen Rolle in einem ansonsten überwiegend aus HTML bestehenden Web. Heute ist es riesig, und einige Websites werden fast vollständig mit JavaScript erstellt. Um JavaScript modularer, sauberer und wartbarer zu machen, führte ES6 eine Möglichkeit ein, Code einfach zwischen JavaScript-Dateien zu teilen. Dabei werden Teile einer Datei exportiert, um sie in einer oder mehreren anderen Dateien zu verwenden, und die benötigten Teile werden dort importiert, wo Sie sie brauchen. Um diese Funktionalität zu nutzen, müssen Sie in Ihrem HTML-Dokument ein Skript mit dem `type` vom Typ `module` erstellen. Hier ist ein Beispiel:

```html
<script type="module" src="filename.js"></script>
```

Ein Skript, das diesen `module`-Typ verwendet, kann nun die `import`- und `export`-Funktionen nutzen, die Sie in den kommenden Aufgaben kennenlernen werden.

# --instructions--

Fügen Sie dem HTML-Dokument ein Skript vom Typ `module` hinzu und geben Sie ihm die Quelldatei `index.js`.

# --hints--

Sie sollten ein `script`-Tag erstellen.

```js
assert(__helpers.removeJSComments(code).match(/<\s*script[^>]*>\s*<\/\s*script\s*>/g));
```

Ihr `script`-Tag sollte das Attribut `type` mit dem Wert `module` haben.

```js
assert(
  __helpers.removeJSComments(code).match(
    /<\s*script\s+[^t]*type\s*=\s*('|")module\1[^>]*>\s*<\/\s*script\s*>/g
  )
);
```

Ihr `script`-Tag sollte ein `src` von `index.js` haben.

```js
assert(
  __helpers.removeJSComments(code).match(
    /<\s*script\s+[^s]*src\s*=\s*('|")index\.js\1[^>]*>\s*<\/\s*script\s*>/g
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <!-- Only change code below this line -->

    <!-- Only change code above this line -->
  </body>
</html>
```

# --solutions--

```html
<html>
  <body>
    <script type="module" src="index.js"></script>
  </body>
</html>
```
