---
id: 5cfa550e84205a357704ccb6
title: "Verwenden Sie Destrukturierungszuweisung, um Werte aus Objekten zu extrahieren"
challengeType: 1
forumTopicId: 301216
dashedName: use-destructuring-assignment-to-extract-values-from-objects
---

# --description--

<dfn>Destrukturierungszuweisung</dfn> ist eine spezielle Syntax, die mit ES6 eingeführt wurde, um Werte direkt aus einem Objekt übersichtlich zuzuweisen.

Betrachten Sie den folgenden ES5-Code:

```js
const user = { name: 'John Doe', age: 34 };

const name = user.name;
const age = user.age;
```

`name` hätte den Wert des Strings `John Doe`, und `age` hätte die Zahl `34`.

Hier ist eine äquivalente Zuweisung mit der ES6-Destrukturierungssyntax:

```js
const { name, age } = user;
```

Wiederum hätte `name` den Wert des Strings `John Doe`, und `age` die Zahl `34`.

Hier werden die Variablen `name` und `age` erstellt und mit den jeweiligen Werten aus dem `user`-Objekt zugewiesen. Sie sehen, wie viel übersichtlicher das ist.

Sie können so viele oder so wenige Werte aus dem Objekt extrahieren, wie Sie möchten.

# --instructions--

Ersetzen Sie die beiden Zuweisungen durch eine äquivalente Destrukturierungszuweisung. Dabei sollen die Variablen `today` und `tomorrow` weiterhin die Werte von `today` und `tomorrow` aus dem `HIGH_TEMPERATURES`-Objekt zugewiesen bekommen.

# --hints--

Sie sollten die ES5-Zuweisungssyntax entfernen.

```js
assert(
  !__helpers.removeJSComments(code).match(/today\s*=\s*HIGH_TEMPERATURES\.(today|tomorrow)/g)
);
```

Sie sollten Destrukturierung verwenden, um die Variable `today` zu erstellen.

```js
assert(
  __helpers.removeJSComments(code).match(/(var|let|const)\s*{\s*(today[^}]*|[^,]*,\s*today)\s*}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g)
);
```

Sie sollten Destrukturierung verwenden, um die Variable `tomorrow` zu erstellen.

```js
assert(
  __helpers.removeJSComments(code).match(/(var|let|const)\s*{\s*(tomorrow[^}]*|[^,]*,\s*tomorrow)\s*}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g)
);
```

`today` sollte gleich `77` sein und `tomorrow` sollte gleich `80` sein.

```js
assert(today === 77 && tomorrow === 80);
```

# --seed--

## --seed-contents--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

// Only change code below this line

const today = HIGH_TEMPERATURES.today;
const tomorrow = HIGH_TEMPERATURES.tomorrow;

// Only change code above this line
```

# --solutions--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

const { today, tomorrow } = HIGH_TEMPERATURES;
```
