---
id: 5e44414f903586ffb414c950
title: Побудова проєкту калькулятора ймовірності
challengeType: 23
saveSubmissionToDB: true
forumTopicId: 462364
dashedName: build-a-probability-calculator-project
---

# --description--

Припустімо, що є капелюх, у якому знаходиться 5 синіх кульок, 4 червоні кульки та 2 зелені кульки. Яка ймовірність того, що випадковий вибір 4 кульок міститиме принаймні 1 червону кульку та 2 зелені кульки? Хоча можна обчислити ймовірність за допомогою складної математики, простіше написати програму, яка виконає велику кількість експериментів, щоб оцінити приблизну ймовірність.

Для цього проєкту ви напишете програму, щоб визначити приблизну ймовірність випадкового вибору певних кульок із капелюха.

Спочатку створіть клас `Hat` у `main.py`. Клас має приймати змінну кількість аргументів, які вказують кількість кульок кожного кольору в капелюсі. Наприклад, об’єкт класу можна створити будь-яким із таких способів:

```py
hat1 = Hat(yellow=3, blue=2, green=6)
hat2 = Hat(red=5, orange=4)
hat3 = Hat(red=5, orange=4, black=1, blue=0, pink=2, striped=9)
```

Капелюх завжди створюватиметься з принаймні однією кулькою. Аргументи, передані об’єкту капелюха під час створення, слід перетворити на змінну екземпляра `contents`. `contents` має бути списком рядків, що містить по одному елементу для кожної кульки в капелюсі. Кожен елемент списку має бути назвою кольору, що представляє одну кульку цього кольору. Наприклад, якщо ваш капелюх — `{'red': 2, 'blue': 1}`, то `contents` має бути `['red', 'red', 'blue']`.

Клас `Hat` має мати метод `draw`, який приймає аргумент, що вказує кількість кульок для вибору з капелюха. Цей метод має випадковим чином видаляти кульки з `contents` і повертати ці кульки у вигляді списку рядків. Кульки не повинні повертатися назад у капелюх під час вибору, подібно до експерименту з урною без повернення. Якщо кількість кульок для вибору перевищує доступну кількість, поверніть усі кульки.

Далі створіть функцію `experiment` у `main.py` (не всередині класу `Hat`). Ця функція має приймати такі аргументи:

- `hat`: об’єкт капелюха, що містить кульки, який слід скопіювати всередині функції.
- `expected_balls`: об’єкт, що вказує точну групу кульок, які потрібно спробувати витягти з капелюха для експерименту. Наприклад, щоб визначити ймовірність витягнути 2 сині кульки та 1 червону кульку з капелюха, встановіть `expected_balls` у `{'blue':2, 'red':1}`.
- `num_balls_drawn`: кількість кульок, які потрібно витягти з капелюха в кожному експерименті.
- `num_experiments`: кількість експериментів для виконання. (Чим більше експериментів, тим точнішою буде приблизна ймовірність.)

Функція `experiment` має повертати ймовірність.

Наприклад, якщо ви хочете визначити ймовірність отримати принаймні дві червоні кульки та одну зелену кульку, коли ви витягуєте п’ять кульок із капелюха, що містить шість чорних, чотири червоні та три зелені. Для цього ви виконаєте `N` експериментів, порахуєте, скільки разів `M` ви отримаєте принаймні дві червоні кульки та одну зелену кульку, і оціните ймовірність як `M/N`. Кожен експеримент полягає в тому, щоб почати з капелюха з вказаними кульками, витягнути кілька кульок і перевірити, чи отримали ви кульки, які намагалися витягти.

Ось як ви викликали б функцію `experiment` на основі наведеного прикладу з 2000 експериментів:

```py
hat = Hat(black=6, red=4, green=3)
probability = experiment(hat=hat,
                  expected_balls={'red':2,'green':1},
                  num_balls_drawn=5,
                  num_experiments=2000)
```

Вихідні дані будуть приблизно такими:

```bash
0.356
```

Оскільки це базується на випадкових виборах, ймовірність буде трохи відрізнятися щоразу, коли код запускається.

_Підказка: Розгляньте використання модулів, які вже імпортовані на початку. Не ініціалізуйте випадкове зерно всередині файлу._

Примітка: відкрийте консоль браузера за допомогою F12, щоб побачити більш докладний вивід тестів.

# --hints--

Створення об’єкта `hat` має додати правильний вміст.

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_class_contents(self):
        hat = probability_calculator.Hat(red=3,blue=2)
        actual = hat.contents
        expected = ["red","red","red","blue","blue"]
        self.assertEqual(actual, expected, 'Expected creation of hat object to add correct contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

Метод `draw` у класі `hat` має зменшувати кількість елементів у вмісті.


```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)

class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_draw(self):
        hat = probability_calculator.Hat(red=5,blue=2)
        actual = hat.draw(2)
        expected = ['blue', 'red']
        self.assertEqual(actual, expected, 'Expected hat draw to return two random items from hat contents.')
        actual = len(hat.contents)
        expected = 5
        self.assertEqual(actual, expected, 'Expected hat draw to reduce number of items in contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

Метод `draw` має працювати правильно, коли кількість кульок для вилучення більша за кількість кульок у капелюсі.

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_draw_2(self):
        hat = probability_calculator.Hat(yellow=5,red=1,green=3,blue=9,test=1)
        actual = sorted(hat.draw(20))
        expected = sorted(['yellow', 'yellow', 'yellow', 'yellow', 'yellow', 'red', 'green', 'green', 'green', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'test'])
        self.assertEqual(actual, expected, 'Expected hat draw to return all items from hat contents.')
        actual = len(hat.contents)
        expected = 0
        self.assertEqual(actual, expected, 'Expected hat draw to leave no items in contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

Метод `experiment` має повертати різну ймовірність.


```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_prob_experiment(self):
        hat = probability_calculator.Hat(blue=3,red=2,green=6)
        probability = probability_calculator.experiment(hat=hat, expected_balls={"blue":2,"green":1}, num_balls_drawn=4, num_experiments=1000)
        actual = probability
        expected = 0.272
        self.assertAlmostEqual(actual, expected, delta = 0.01, msg = 'Expected experiment method to return a different probability.')
        hat = probability_calculator.Hat(yellow=5,red=1,green=3,blue=9,test=1)
        probability = probability_calculator.experiment(hat=hat, expected_balls={"yellow":2,"blue":3,"test":1}, num_balls_drawn=20, num_experiments=100)
        actual = probability
        expected = 1.0
        self.assertAlmostEqual(actual, expected, delta = 0.01, msg = 'Expected experiment method to return a different probability.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

# --seed--

## --seed-contents--

```py
import copy
import random

class Hat:
    pass

def experiment(hat, expected_balls, num_balls_drawn, num_experiments):
    pass
```

# --solutions--

```py
import copy
import random

class Hat:
    def __init__(self, **hat):
        self.hat = hat
        contents = []
        for i in hat:
            for j in range(hat[i]):
                contents.append(i)           
        self.contents = contents
               

    def draw(self, number): 
        drawn = []
        if number >= len(self.contents):
            drawn.extend(self.contents)
            self.contents = []
        else:
            for i in range(number):
                drawn.append(
                    self.contents.pop(random.randrange(len(self.contents)))
                )                
        return drawn

def experiment(hat, expected_balls, num_balls_drawn, num_experiments):    
    expected_balls_list = []
    drawn_list = []
    success = 0
    for i in expected_balls:
        for j in range(expected_balls[i]):
            expected_balls_list.append(i)
    for j in range(num_experiments):
        hat_copy = copy.deepcopy(hat)
        drawn_list.append(hat_copy.draw(num_balls_drawn))        
        exp_ball_list_copy = expected_balls_list[:]
        for k in range(len(drawn_list[j])):
            try:
                ind = exp_ball_list_copy.index(drawn_list[j][k])
                exp_ball_list_copy.pop(ind)
            except:
                continue
        if len(exp_ball_list_copy) == 0:
            success += 1        
    probability = success/num_experiments    
    return probability
```
