---
id: 66da326c02141df538f29ba5
title: 결제 페이지 만들기
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 텍스트가 `h1`인 `Checkout` 요소가 있어야 합니다.
1. `section` 요소 바로 뒤에 `h1` 요소 두 개가 있어야 합니다.
1. 첫 번째 섹션 내에 텍스트가 `h2`인 `Your Cart` 요소가 있어야 합니다.
1. 첫 번째 섹션에 적절한 대체 텍스트가 있는 상품 이미지가 있어야 합니다. 원한다면 이 이미지를 사용할 수 있습니다: `https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. 두 번째 섹션 내에 텍스트가 `h2`인 `Payment Information` 요소가 있어야 합니다.
1. 두 번째 섹션 내에 `form` 요소가 있어야 합니다.
1. 폼 내에 `id`과 `name`이 `card-name`인 입력과 `type`이 `text`인 입력, 그리고 이와 연결된 `label`가 있어야 합니다.
1. 폼 내에 `id`과 `name`이 `card-number`인 입력과 `type`가 `text`인 입력, 그리고 이와 연결된 `label`가 있어야 합니다.
1. `input` 속성이 있는 `required` 요소가 최소 두 개 있어야 합니다.
1. 필수 입력란이 시각적으로 표시되도록 각 필수 입력에 대해 `span` 요소 안에 텍스트가 `*`이고 `aria-hidden`이 `true`인 `label` 요소를 포함해야 합니다.
1. 카드 번호 입력 바로 뒤에 카드 번호 형식을 설명하는 도움말 텍스트가 있는 `p` 요소가 있어야 합니다. `p`은 `id`이 `card-number-help`이고 카드 번호 입력이 `aria-describedby`을 사용해 참조해야 합니다.

# --hints--

텍스트가 `h1`인 `Checkout` 요소가 있어야 합니다.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

페이지에 `h1` 요소는 하나만 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

`section` 요소 뒤에 `h1` 요소가 최소 두 개 있어야 합니다.

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

첫 번째 섹션 내에 텍스트가 `h2`인 `Your Cart` 요소가 있어야 합니다.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

첫 번째 섹션에 대체 텍스트가 있는 이미지가 최소 하나 있어야 합니다.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

두 번째 섹션 내에 텍스트가 `h2`인 `Payment Information` 요소가 있어야 합니다.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

두 번째 `form` 요소 내에 `section` 요소가 있어야 합니다.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

폼 내에 `input`인 `id`과 `name`이 있는 `card-name` 요소가 있어야 합니다.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

카드 이름 입력에는 `type`인 `text`가 있어야 합니다.

```js
const cardNameInput = document.getElementById('card-name');
assert.exists(cardNameInput);
assert.equal(cardNameInput.getAttribute('type'), 'text');
```

폼 내에 `input`인 `id`와 `name`이 있는 `card-number` 요소가 있어야 합니다.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

카드 번호 입력에는 `type`인 `text`가 있어야 합니다.

```js
const cardNumberInput = document.getElementById('card-number');
assert.exists(cardNumberInput);
assert.equal(cardNumberInput.getAttribute('type'), 'text');
```

`input`의 `type`가 아닌 모든 `submit` 요소에는 `label` 요소가 연결되어 있어야 합니다.

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
});
```

`input` 속성이 있는 `required` 요소가 최소 두 개 있어야 합니다.

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required);
assert.isAtLeast(requiredInputs?.length, 2)
```

각 필수 입력에 대해 `span` 요소 안에 텍스트가 `*`인 `label` 요소를 포함해야 합니다.

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.textContent.trim(), '*');
});
```

`span` 요소에는 `aria-hidden`이 `true`로 설정되어 있어야 합니다.

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.getAttribute('aria-hidden'), 'true');
});
```

카드 번호 입력 바로 뒤에 `p`인 `id`이 있는 `card-number-help` 요소가 있어야 합니다.

```js
const cardNumberInput = document.getElementById('card-number');
const cardNumberHelp = document.getElementById('card-number-help');

assert.exists(cardNumberHelp);
const nextElement = cardNumberInput.parentElement?.tagName === 'LABEL'
  ? cardNumberInput.parentElement.nextElementSibling
  : cardNumberInput.nextElementSibling;
assert.equal(nextElement, cardNumberHelp);
```

카드 번호 도움말 텍스트는 비어 있으면 안 됩니다.

```js
const cardNumberHelp = document.getElementById('card-number-help');
assert.isString(cardNumberHelp?.textContent);
assert.isAbove(cardNumberHelp?.textContent.trim().length, 0);
```

카드 번호 입력에는 `aria-describedby`가 `card-number-help`로 설정되어 있어야 합니다.

```js
const cardNumberInput = document.getElementById('card-number');
assert.equal(cardNumberInput.getAttribute('aria-describedby'), 'card-number-help');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-name" name="card-name" required>
                </div>
                <div>
                    <label for="card-number">Card Number <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-number" name="card-number" aria-describedby="card-number-help" required>
                    <p id="card-number-help">Please enter your 16-digit card number without spaces or dashes.</p>
                </div>
                <div>
                    <label for="expiry-date">Expiry Date <span aria-hidden="true">*</span></label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required>
                </div>
                <div>
                    <label for="cvv">CVV <span aria-hidden="true">*</span></label>
                    <input type="text" id="cvv" name="cvv" required aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```
