---
id: 0095d0567b6817114989a1dd
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop implementerai l'algoritmo del percorso più breve in JavaScript. Scriverai una funzione che calcola il percorso più breve tra i nodi in un grafo pesato e restituisce sia le distanze sia i percorsi seguiti.

Per esempio, dato un grafo in cui le città sono collegate da strade con distanze diverse, l'algoritmo troverà la rotta più breve da una città all'altra. Se vuoi viaggiare dalla città A alla città D, l'algoritmo potrebbe trovare che andare A ⇨ B ⇨ C ⇨ D (totale: 15 km) è più breve che andare direttamente A ⇨ D (20 km).

Per iniziare, definisci una variabile chiamata `INF` e assegnale il valore `Infinity`. Più avanti la userai per indicare che non esiste una connessione diretta tra due nodi.

# --hints--

Dovresti definire una variabile chiamata `INF`.

```js
assert.isDefined(INF);
```

La tua variabile `INF` dovrebbe avere il valore `Infinity`.

```js
assert.strictEqual(INF, Infinity);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
