---
id: 69c918841987bf16c1d1baca
title: Passaggio 19
challengeType: 1
dashedName: step-19
---

# --description--

Quando trovi un percorso più breve verso un vicino, devi anche aggiornare il percorso registrato per raggiungerlo.

In JavaScript, puoi creare un nuovo array combinando un array esistente e un nuovo elemento usando l'operatore di espansione:

```js
const newPath = [...existingArray, newElement];
```

All'interno del blocco annidato `if`, aggiorna `paths[nodeNo]` in modo che sia il percorso verso il nodo corrente con `nodeNo` aggiunto in coda usando la sintassi di espansione come segue: `[...paths[current], nodeNo]`.

# --hints--

Dovresti aggiornare `paths[nodeNo]` all'interno del blocco annidato `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*\[\s*\w+\s*\]\s*=/
);
```

`paths[nodeNo]` dovrebbe essere impostato su `[...paths[current], nodeNo]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*\[\s*\w+\s*\]\s*=\s*\[\s*\.\.\.\s*paths\s*\[\s*current\s*\]\s*,\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          --fcc-editable-region--

          --fcc-editable-region--
        }
      }
    }
  }
}
```
