---
id: 69c918841987bf16c1d1baba
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Ti servirà un array 2D per rappresentare la matrice di adiacenza del grafo. Ogni valore rappresenta il peso (distanza) del collegamento tra due nodi. Un valore `INF` indica che non esiste un collegamento diretto.

Crea un'altra variabile chiamata `adjMatrix` e assegnale un array 2D con i seguenti valori:

```js
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

Dovresti creare una variabile chiamata `adjMatrix`.

```js
assert.isDefined(adjMatrix);
```

`adjMatrix` dovrebbe essere un array 2D con 6 righe e 6 colonne.

```js
assert.isArray(adjMatrix);
assert.lengthOf(adjMatrix, 6);
adjMatrix.forEach(row => assert.lengthOf(row, 6));
```

`adjMatrix` dovrebbe contenere i valori corretti.

```js
assert.deepEqual(adjMatrix[0], [0, 5, 3, INF, 11, INF]);
assert.deepEqual(adjMatrix[1], [5, 0, 1, INF, INF, 2]);
assert.deepEqual(adjMatrix[2], [3, 1, 0, 1, 5, INF]);
assert.deepEqual(adjMatrix[3], [INF, INF, 1, 0, 9, 3]);
assert.deepEqual(adjMatrix[4], [11, INF, 5, 9, 0, INF]);
assert.deepEqual(adjMatrix[5], [INF, 2, INF, 3, INF, 0]);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
--fcc-editable-region--

--fcc-editable-region--
```
