---
id: 69c918841987bf16c1d1bac7
title: Passaggio 16
challengeType: 1
dashedName: step-16
---

# --description--

Ora che il nodo corrente è segnato come visitato, devi esaminare tutti i suoi vicini per vedere se puoi raggiungerli in modo più efficiente passando per il nodo corrente.

Dopo `visited[current] = true`, aggiungi un loop `for` sui vicini che itera su ogni `nodeNo` da `0` a `n`. Nota che questo è un altro loop interno allo stesso livello del primo.

All’interno del loop, dichiara una variabile chiamata `distance` e assegnale `matrix[current][nodeNo]`. Questo ti dà il peso del bordo tra il nodo corrente e il suo vicino.

# --hints--

Dovresti aggiungere un loop `for` sui vicini dopo aver segnato il nodo corrente come visitato.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(/
);
```

Il loop `for` sui vicini dovrebbe iterare da `0` a `n`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(\s*(var|let)\s+\w+\s*=\s*0\s*;\s*\w+\s*<\s*n\s*;\s*\w+\+\+\s*\)/
);
```

Dovresti avere una variabile `distance` assegnata a `matrix[current][nodeNo]` all’interno del loop sui vicini.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distance\s*=\s*matrix\s*\[\s*current\s*\]\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
