---
id: 69c918841987bf16c1d1bacb
title: Passaggio 20
challengeType: 1
dashedName: step-20
---

# --description--

Una volta che il loop principale termina, devi decidere per quale nodo mostrare i risultati.

Se è stato fornito un `targetNode` specifico, mostra i risultati solo per quel nodo. Altrimenti, mostra i risultati per tutti i nodi.

In JavaScript, puoi ottenere tutti gli indici di un array come iterabile usando `[...Array(n).keys()]`.

Dopo il loop esterno `for`, crea una variabile chiamata `targets`. Usando un'espressione ternaria, se `targetNode` è strettamente diverso da `null`, assegna `[targetNode]`, altrimenti assegna `[...Array(n).keys()]`.

# --hints--

Dovresti creare una variabile `targets` dopo il loop principale.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+targets\s*=/
);
```

La tua variabile `targets` dovrebbe usare un ternario che controlla `targetNode !== null`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /targets\s*=\s*targetNode\s*!==\s*null\s*\?/
);
```

Quando viene fornito un obiettivo, `targets` dovrebbe essere `[targetNode]`, altrimenti dovrebbe essere `[...Array(n).keys()]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /\?\s*\[\s*targetNode\s*\]\s*:\s*\[\s*\.\.\.\s*Array\s*\(\s*n\s*\)\s*\.keys\s*\(\s*\)\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
