---
id: 69c918841987bf16c1d1bacd
title: Passaggio 22
challengeType: 1
dashedName: step-22
---

# --description--

Vuoi mostrare i risultati solo per i nodi raggiungibili e diversi dal nodo di partenza.

All’interno del loop `for...of`, aggiungi un’istruzione `if` che verifichi se `nodeNo` è strettamente uguale a `startNode` **oppure** `distances[nodeNo]` è strettamente uguale a `INF`. Se una delle due condizioni è vera, usa `continue` all’interno del blocco `if` per passare all’iterazione successiva.

# --hints--

Dovresti avere un’istruzione `if` con la condizione `nodeNo === startNode || distances[nodeNo] === INF` all’interno del tuo loop `for...of`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*\w+\s*===\s*startNode\s*\|\|\s*distances\s*\[\s*\w+\s*\]\s*===\s*INF\s*\)/
);
```

Dovresti usare `continue` nel blocco `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*\w+\s*===\s*startNode\s*\|\|\s*distances\s*\[\s*\w+\s*\]\s*===\s*INF\s*\)[\s\S]*?continue/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }

  const targets = targetNode !== null ? [targetNode] : [...Array(n).keys()];
--fcc-editable-region--
  for (const nodeNo of targets) {

  }
--fcc-editable-region--
}
```
