---
id: 69c918841987bf16c1d1bac4
title: Passaggio 13
challengeType: 1
dashedName: step-13
---

# --description--

Se la condizione è vera, il nodo corrente è il miglior candidato non visitato trovato finora. Devi aggiornare le tue variabili di tracciamento per riflettere questo.

All'interno dell'istruzione `if`, aggiorna `minDistance` a `distances[nodeNo]` e `current` a `nodeNo`.

# --hints--

Dovresti impostare `minDistance` su `distances[nodeNo]` all'interno del blocco `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /minDistance\s*=\s*distances\s*\[\s*nodeNo\s*\]/
);
```

Dovresti impostare `current` su `nodeNo` all'interno del blocco `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /current\s*=\s*nodeNo/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        --fcc-editable-region--

        --fcc-editable-region--
      }
    }
  }
}
```
