---
id: 69c918841987bf16c1d1babd
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Devi tenere traccia della distanza più breve conosciuta dal nodo di partenza a ogni altro nodo. Per cominciare, assumerai che ogni nodo sia a distanza infinita.

In JavaScript, puoi creare un array precompilato con un valore usando `new Array(n).fill(value)`. Per esempio, `new Array(3).fill(0)` crea `[0, 0, 0]`.

All'interno della funzione `shortestPath`, crea una variabile chiamata `distances` e assegnale `new Array(n).fill(INF)`.

# --hints--

Dovresti creare una variabile chiamata `distances` all'interno di `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distances\s*=/
);
```

La tua variabile `distances` dovrebbe essere impostata su `new Array(n).fill(INF)`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*INF\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;

}
--fcc-editable-region--
```
