---
id: 69c918841987bf16c1d1babf
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Devi anche tenere traccia del percorso effettivo seguito per raggiungere ogni nodo. Archivierai una lista di percorsi dove ogni voce è un array di indici dei nodi che rappresentano la rotta seguita.

Inizialmente, il percorso di ogni nodo conterrà solo se stesso. Puoi creare questa struttura usando `Array.from()`:

```js
Array.from({ length: n }, (_, i) => [i]);
```

Questo crea un array di `n` elementi, dove l'elemento all'indice `i` è `[i]`.

Sempre all'interno della funzione `shortestPath`, crea una variabile chiamata `paths` e assegnale `Array.from({ length: n }, (_, i) => [i])`.

# --hints--

Dovresti creare una variabile `paths` all'interno della funzione `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+paths\s*=/
);
```

La tua variabile `paths` dovrebbe essere inizializzata usando `Array.from`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*=\s*Array\.from/
);
```

La tua variabile `paths` dovrebbe usare `{ length: n }` come primo argomento di `Array.from`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\(\s*\{\s*length\s*:\s*n\s*\}/
);
```

Il secondo argomento di `Array.from` dovrebbe essere una callback `(_, i) => [i]`, dove `_` è un segnaposto per l'elemento non usato e `i` è l'indice. Dovrebbe restituire `[i]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\([\s\S]*?,\s*function\s*\(\s*\w+\s*,\s*(\w+)\s*\)\s*\{[\s\S]*?return\s*\[\s*\1\s*\]|Array\.from\s*\([\s\S]*?,\s*\(\s*\w+\s*,\s*\w+\s*\)\s*=>\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
