---
id: 69c937c73f23b1140bcd297a
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Il parametro `targetNode` Ă¨ opzionale. Quando non viene fornito, la funzione calcolerĂ  i percorsi piĂ¹ brevi da `startNode` a tutti gli altri nodi nel grafo.

Imposta il parametro `targetNode` con un valore predefinito di `null`.

# --hints--

Il parametro `targetNode` della tua funzione `shortestPath` dovrebbe avere come valore predefinito `null`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /targetNode\s*=\s*null/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode) {

}
--fcc-editable-region--
```
