---
id: 69c918841987bf16c1d1bac1
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

Ora aggiungerai il loop principale che guida l'algoritmo. Viene eseguito una volta per ogni nodo nel grafo, selezionando ogni volta il nodo non visitato più vicino.

All'interno della funzione `shortestPath`, aggiungi un loop `for` che viene eseguito `n` volte usando `i` come variabile del loop.

All'inizio del corpo del loop, usando `let`, dichiara due variabili `minDistance` e `current` impostate rispettivamente su `INF` e `-1`. Queste terranno traccia del nodo non visitato più vicino trovato in ogni iterazione.

# --hints--

Dovresti avere un loop `for` all'interno di `shortestPath` che viene eseguito `n` volte. Non dimenticare di impostare `i` con `let` al suo interno.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+i\s*=\s*0\s*;\s*i\s*<\s*n\s*;\s*i\+\+\s*\)/
);
```

Dovresti avere una variabile `minDistance` impostata su `INF` all'interno del loop.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+minDistance\s*=\s*INF/
);
```

Dovresti avere una variabile `current` impostata su `-1` all'interno del loop.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+current\s*=\s*-1/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
--fcc-editable-region--

--fcc-editable-region--
}
```
