---
id: 69c918841987bf16c1d1bac8
title: Passaggio 17
challengeType: 1
dashedName: step-17
---

# --description--

Prima di aggiornare le distanze, devi verificare che il vicino valga la pena di essere considerato. Deve esserci un arco reale verso di esso (`distance !== INF`) e non deve essere stato ancora visitato.

All’interno del loop `for` dei vicini, aggiungi un’istruzione `if` che controlli se `distance` non è strettamente uguale a `INF` **e** se `nodeNo` non è visitato.

All’interno del blocco `if`, dichiara una variabile `newDistance` assegnata a `distances[current] + distance`. Questo è il costo totale per raggiungere il vicino passando per il nodo corrente.

# --hints--

Dovresti aggiungere un’istruzione `if` con la condizione `distance !== INF && !visited[nodeNo]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*distance\s*!==\s*INF\s*&&\s*!\s*visited\s*\[\s*\w+\s*\]\s*\)/
);
```

Dovresti dichiarare una variabile `newDistance` con il valore `distances[current] + distance` all’interno del blocco `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+newDistance\s*=\s*distances\s*\[\s*current\s*\]\s*\+\s*distance/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
