---
id: 69c918841987bf16c1d1bac0
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

Mentre l'algoritmo viene eseguito, devi tenere traccia dei nodi già processati per non visitarli di nuovo.

Crea una variabile chiamata `visited` e assegnale `new Array(n).fill(false)`.

# --hints--

Dovresti creare una variabile chiamata `visited`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+visited\s*=/
);
```

La tua variabile `visited` dovrebbe essere impostata su `new Array(n).fill(false)`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*false\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
--fcc-editable-region--

--fcc-editable-region--
}
```
