---
id: 69c918841987bf16c1d1bace
title: Passaggio 23
challengeType: 1
dashedName: step-23
---

# --description--

Ora che sai che il nodo è raggiungibile e non è l'inizio, formatta il suo percorso in una stringa leggibile.

L'array `paths[nodeNo]` contiene la sequenza degli indici dei nodi visitati per raggiungere `nodeNo`. Puoi trasformarla in una stringa leggibile come `0 -> 2 -> 3` usando il metodo `.join()`.

All'interno del loop `for...of`, dopo l'istruzione `if` con `continue`, dichiara una variabile chiamata `path` e assegnale il valore `paths[nodeNo].join(' -> ')`.

# --hints--

Dovresti avere una variabile chiamata `path` all'interno del loop `for...of`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /for\s*\(\s*(const|let)\s+nodeNo\s+of\s+targets\s*\)[\s\S]*?(const|let)\s+path\s*=/
);
```

La tua variabile `path` dovrebbe avere il valore `paths[nodeNo].join(' -> ')`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /path\s*=\s*paths\s*\[\s*\w+\s*\]\s*\.\s*join\s*\(\s*['"]\s*->\s*['"]\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }

  const targets = targetNode !== null ? [targetNode] : [...Array(n).keys()];
  for (const nodeNo of targets) {
    if (nodeNo === startNode || distances[nodeNo] === INF) {
      continue;
    }
    --fcc-editable-region--

    --fcc-editable-region--
  }
}
```
