---
id: 691f7bbf76229bb97c60827d
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop, eserciterai l'uso dei loop `for...of` costruendo una funzione che conta quante volte una stringa appare in un array di stringhe.

Prima di scrivere quella funzione, creerai una più semplice che stampa ogni carattere di una stringa nella console.

Inizia definendo una funzione vuota chiamata `printCharacters` con il parametro `str`.

# --hints--

Dovresti creare una funzione `printCharacters`.

```js
assert.isFunction(printCharacters);
```

La tua funzione `printCharacters` dovrebbe avere un parametro chiamato `str`.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.printCharacters;
assert.lengthOf(parameters, 1);
assert.equal(parameters[0].toString(), 'str');
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
