---
id: 6920c9cda1c7f056f75e4b1a
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Ora è il momento di scrivere il tuo primo loop `for...of`. Nelle lezioni precedenti, hai imparato che puoi usare un loop `for...of` per iterare su valori provenienti da un iterabile (ad esempio, stringhe o array).

Ecco un esempio di loop `for...of`:

```js
for (const num of [1, 2, 3]) {
  // code block to be executed
}
```

Aggiungi un loop `for...of` con un blocco di codice vuoto all’interno della tua funzione. Dovrebbe iterare su ogni carattere nell’argomento `str`.

# --hints--

La tua funzione `printCharacters` dovrebbe avere un loop `for...of` all’interno del suo blocco di codice.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functionprintCharacters\(str\)\{for\(\S+of\S+\)\{\S*?\}\}/);
```

Il tuo loop `for...of` dovrebbe iterare sull’argomento `str` della tua funzione.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function printCharacters\(str\) \{\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+str/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
