---
id: 6920d189e36b6b60a6330fee
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Ora, all'interno del loop, registra la variabile `char` nella console.

# --hints--

Dovresti avere un'istruzione `console.log` con `char` come argomento all'interno del loop.

```js
const explorer = await __helpers.Explorer(code);
const normalizedFunctionCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(explorer.allFunctions.printCharacters?.toString()));
const loopCodeBlock = normalizedFunctionCode.match(/for\(constcharofstr\)\{(\S*?)}/)?.[1];
assert.match(loopCodeBlock, /console\.log\(char\)/);
```

Quando la funzione `printCharacters` viene chiamata con l'argomento `'test'`, l'istruzione `console.log()` all'interno della funzione dovrebbe essere chiamata quattro volte con i seguenti argomenti, in ordine: `'t'`, `'e'`, `'s'` e `'t'`. Invece, è stata chiamata con i seguenti argomenti: `--fcc-actual--`.

```js
const consoleLogSpy = __helpers.spyOn(console, 'log');
printCharacters('test');
assert.deepEqual(consoleLogSpy.calls, [['t'], ['e'], ['s'], ['t']]);
consoleLogSpy.restore();
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
