---
id: 6921d688b8d58581d34459b9
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Nei prossimi passaggi costruirai una funzione che conta quante volte una stringa appare in un array di stringhe.

Per cominciare, definisci una funzione vuota chiamata `getMatchedWordCount` con i parametri `sentence` e `match` in questo ordine.

# --hints--

Dovresti creare una funzione `getMatchedWordCount`.

```js
assert.isFunction(getMatchedWordCount);
```

La tua funzione `getMatchedWordCount` dovrebbe avere due parametri: `sentence` e `match`, in questo ordine.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
