---
id: 6921d90b5ce50185505d2004
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Poiché la funzione `getMatchedWordCount` deve restituire un conteggio numerico, inizierai dichiarando una variabile contatore.

Crea una variabile chiamata `count` usando la parola chiave `let`, inizializzala con il valore `0` e aggiungi un'istruzione `return` che restituisca la variabile `count`.

# --hints--

Dovresti dichiarare la variabile `count` all'interno della funzione `getMatchedWordCount`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /(const|let|var)count/);
```

Per permettere che la variabile `count` venga incrementata in seguito, dichiarala con la parola chiave `let`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /letcount/);
```

La tua funzione dovrebbe contenere un'assegnazione del valore `0` alla variabile `count`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /count=0/);
```

La tua funzione dovrebbe contenere un'istruzione `return` che restituisca la variabile `count`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /returncount/);
```

La tua funzione `getMatchedWordCount` dovrebbe restituire il valore `0`.

```js
assert.strictEqual(getMatchedWordCount(["foo", "bar"], "foo"), 0);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
