---
id: 6921de8d2a111d0341cff702
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Ottimo, la tua funzione restituisce un `0`, ma ancora non conta nulla.

Per sistemarlo, all'interno di `getMatchedWordCount` crea un loop `for...of` con un blocco di codice vuoto che itera su ogni parola in `sentence`.

# --hints--

La tua funzione `getMatchedWordCount` dovrebbe avere un loop `for...of` all'interno del suo blocco di codice.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functiongetMatchedWordCount\(sentence,match\)\{letcount=0\;for\(\S+of\S+\)\{\S*?\}returncount;\}/);
```

Il tuo loop `for...of` dovrebbe iterare sull'argomento `sentence` della tua funzione.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function getMatchedWordCount\(sentence, match\) \{\s*let count = 0\;\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*return count\;\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+sentence/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
--fcc-editable-region--
  
--fcc-editable-region--
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
