---
id: 6921d90b5ce50185505d2004
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Poiché la funzione `getMatchedWordCount` deve restituire un conteggio numerico, inizierai dichiarando una variabile contatore.

Crea una variabile chiamata `count` usando la parola chiave `let`, inizializzala con il valore `0` e aggiungi un'istruzione `return` che restituisca la variabile `count`.

# --hints--

Dovresti dichiarare la variabile `count` all'interno della funzione `getMatchedWordCount`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getMatchedWordCount?.variables.count);
```

Per permettere che la variabile `count` venga incrementata in seguito, dichiarala con la parola chiave `let`.

```js
const explorer = await __helpers.Explorer(code);
const countVar = explorer.allFunctions.getMatchedWordCount?.variables.count;
assert.isTrue(countVar?.toString().startsWith('let'));
```

La tua funzione dovrebbe contenere un'assegnazione del valore `0` alla variabile `count`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getMatchedWordCount?.variables.count?.value.matches('0'));
```

La tua funzione dovrebbe contenere un'istruzione `return` che restituisca la variabile `count`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getMatchedWordCount?.hasReturn('count'));
```

La tua funzione `getMatchedWordCount` dovrebbe restituire il valore `0`.

```js
assert.strictEqual(getMatchedWordCount(["foo", "bar"], "foo"), 0);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
