---
id: 6921dae10700e300f26994e3
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Per vedere l’output della tua funzione `getMatchedWordCount` nel suo stato attuale, aggiungi una istruzione `console.log` subito sotto. All’interno di `console.log()`, chiama la funzione con l’argomento `["I", "really", "really", "really", "like", "to", "code"]` per il parametro `sentence` e `"really"` per il parametro `match`.

# --hints--

Dovresti chiamare la funzione `getMatchedWordCount` all’interno di una istruzione `console.log`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
assert.match(normalizedCode, /console\.log\(getMatchedWordCount\(\S*?\)\)/);
```

Dovresti chiamare la funzione `getMatchedWordCount` con l’argomento `["I", "really", "really", "really", "like", "to", "code"]` per il parametro `sentence` e `"really"` per il parametro `match`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/)[0];
assert.match(normalizedConsoleLogCode, /\[('|"|`)I\1,('|"|`)really\2,('|"|`)really\3,('|"|`)really\4,('|"|`)like\5,('|"|`)to\6,('|"|`)code\7\],('|"|`)really\8/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
