---
id: 69220a03729c43082ad3e9f6
title: Passaggio 11
challengeType: 1
dashedName: step-11
---

# --description--

Hai finito di lavorare sulla tua funzione `getMatchedWordCount`!

Ora la testerai con dati diversi per vedere come si comporta.

Aggiungi una nuova istruzione `console.log` sotto quella esistente che chiama `getMatchedWordCount` con questi parametri:

- sentence: `["Do", "not", "fear", "the", "dandy", "lion"]`
- match: `"dandy"`

Congratulazioni! Hai completato questo workshop.

# --hints--

Dovresti avere una seconda istruzione `console.log` che chiama la funzione `getMatchedWordCount` al suo interno.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/g)[1];
assert.exists(normalizedConsoleLogCode);
```

Nella tua seconda istruzione `console.log`, chiama la funzione `getMatchedWordCount` con `["Do", "not", "fear", "the", "dandy", "lion"]` come argomento `sentence` e `"dandy"` come argomento `match`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/g)[1];
assert.match(normalizedConsoleLogCode, /\[('|"|`)Do\1,('|"|`)not\2,('|"|`)fear\3,('|"|`)the\4,('|"|`)dandy\5,('|"|`)lion\6\],('|"|`)dandy\7/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
    if (word === match) {
      count++;
    }
    console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }
  
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
    if (word === match) {
      count++;
    }
    console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }
  
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);

console.log(
  getMatchedWordCount(
    ["Do", "not", "fear", "the", "dandy", "lion"],
    "dandy"
  )
);
```
