---
id: 6978bd465213832659f3338b
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

Ora, all’interno del loop, registra nella console il seguente letterale template: `Checking "${word}" against "${match}" | Running count: ${count}`

# --hints--

Dovresti avere esattamente una istruzione `console.log` all’interno del loop nella tua funzione `getMatchedWordCount`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const functionRegex = /(functiongetMatchedWordCount\(sentence,match\)\{\S*?returncount;\})/;
const normalizedFunctionCode = normalizedCode.match(functionRegex)[1];
console.log(normalizedFunctionCode);

assert.match(normalizedFunctionCode, /for\(constwordofsentence\)\{\S*?console\.log\(\S*?\)\S*?\}/);
assert.notMatch(normalizedFunctionCode, /for\(constwordofsentence\)\{\S*?(console\.log\(\S*?\))\S*?(console\.log\(\S*?\))\S*?\}/);
```

La tua istruzione `console.log` all’interno del loop nella funzione `getMatchedWordCount` dovrebbe usare un letterale template con i segnaposto seguenti: `${word}`, `${match}` e `${count}`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const functionRegex = /(functiongetMatchedWordCount\(sentence,match\)\{\S*?returncount;\})/;
const normalizedFunctionCode = normalizedCode.match(functionRegex)[1];
const consoleLogArguments = normalizedFunctionCode.match(/for\(constwordofsentence\)\{\S*?console\.log\((\S*?)\)\S*?\}/)[1]
const templateLiteral = consoleLogArguments.match(/(^\`\S*\`$)/)[1];

assert.include(templateLiteral, "${word}");
assert.include(templateLiteral, "${match}");
assert.include(templateLiteral, "${count}");
```

Quando la funzione `getMatchedWordCount` viene chiamata con `(['foo'], 'bar')`, l’istruzione `console.log()` all’interno della funzione dovrebbe essere chiamata una volta con la stringa `Checking "foo" against "bar" | Running count: 0`. Invece, è stata chiamata con ogni stringa nella seguente lista separata da virgole: `--fcc-actual--`.

```js
const consoleLogSpy = __helpers.spyOn(console, 'log');
getMatchedWordCount(['foo'], 'bar');
assert.deepEqual(consoleLogSpy.calls, [['Checking "foo" against "bar" | Running count: 0']]);
consoleLogSpy.restore();
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
--fcc-editable-region--
    
--fcc-editable-region--
  }

  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
