---
id: 699f49669ea51a31d2ba89d8
title: الخطوة 12
challengeType: 1
dashedName: step-12
---

# --description--

كما تتذكر من الدروس السابقة، يمكنك تحديد نوع القيمة المرجعة لدالة بشكل صريح هكذا:

```ts
function circleArea(radius: number): number {
  return Math.PI * radius * radius;
}
```

في هذا المثال، نوع القيمة المرجعة للدالة `circleArea` هو `number`. إذا حاولت إرجاع نوع مختلف، سيعرض TypeScript رسالة خطأ.

حدّث دالتك `getAverage` لتحدد نوع القيمة المرجعة صراحةً كـ `number`.

# --hints--

يجب أن تكون نوع القيمة المرجعة لدالتك `getAverage` هو `number`.

```js
const explorer = await __helpers.Explorer(code);
const { getAverage } = explorer.allFunctions;
assert.isTrue(getAverage.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]) {
--fcc-editable-region--
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
