---
id: 699f412114c87221f16a99e3
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

ميزة أخرى لاستخدام TypeScript هي أنه يكتشف الأخطاء إذا مررت عددًا كبيرًا جدًا أو قليلًا جدًا من المعلمات في استدعاء دالة.

لتوضيح ذلك، حدّث استدعاء دالة `square` ليأخذ معلمتين هما `5` و `10`.

# --hints--

يجب أن يحتوي استدعاء دالة `square` على معلمتين هما `5` و `10`.

```js
const explorer = await __helpers.Explorer(code);
const { result } = explorer.variables;
assert.isTrue(result.value.matches("square(5, 10)"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5, 10);

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
