---
id: 699f4a3e7b4c6bfa538cc4f7
title: الخطوة 13
challengeType: 1
dashedName: step-13
---

# --description--

إذا فتحت وحدة التحكم، سترى رسالة الخطأ التالية:

```md
Type 'string' is not assignable to type 'number'.
```

لحل هذا الخطأ، عليك إعادة القيمة المرجعة إلى `sum / numbers.length`.

# --hints--

يجب أن تُرجع دالة `getAverage` رقمًا.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

يجب أن تُرجع دالة `getAverage` متوسط الأرقام في المصفوفة.

```js
assert.strictEqual(getAverage([1, 2, 3]), 2);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
--fcc-editable-region--
  return "something";
--fcc-editable-region--
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

