---
id: 699f1d72e906e4616155290d
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستستمر في تعلم التعليقات التوضيحية للنوع وكيفية عملها مع الدوال من خلال بناء مجموعة أدوات رياضية آمنة من حيث النوع.

ابدأ بإنشاء دالة تسمى `square` تقبل معلمة وتُرجع مربع الرقم.

# --hints--

يجب أن يكون لديك دالة تسمى `square`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.square);
```

يجب أن تحتوي دالتك `square` على معلمة واحدة.

```js
const explorer = await __helpers.Explorer(code);
assert.lengthOf(explorer.allFunctions.square.parameters, 1);
```

يجب أن تُرجع دالتك `square` مربع الرقم.

```js
assert.strictEqual(square(2), 4);
assert.strictEqual(square(3), 9);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
