---
id: 699f3738921fb79a651ded56
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

حان الوقت الآن لاستدعاء دالتك.

أنشئ متغيرًا اسمه `result` وقم بتعيين استدعاء الدالة `square` له مع معلمة قيمتها `5`.

أسفل ذلك، سجّل قيمة `result` في وحدة التحكم.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن يكون لديك متغير اسمه `result`.

```js
assert.isDefined(result);
```

يجب أن يكون المتغير `result` رقمًا.

```js
assert.isNumber(result);
```

يجب أن يكون المتغير `result` هو مربع العدد `5`.

```js
assert.strictEqual(result, 25);
```

يجب أن تسجل المتغير `result` في وحدة التحكم.

```js
assert.equal(getLogs()[0], result);
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
