---
id: 699f38937b053e4679878ff1
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

عند العمل مع جافاسكريبت العادية، لا يوجد ما يمنعك من استدعاء الدالة `square` باستخدام سلسلة نصية أو مصفوفة، مما يؤدي إلى نتائج غير متوقعة.

لتوضيح ذلك، غيّر استدعاء الدالة ليستخدم السلسلة النصية `"something"` بدلًا من الرقم `5`.

في الدرس التالي، ستتعلم كيف تستخدم TypeScript لإضافة أمان نوعي إلى شفرتك ومنع هذه الأنواع من الأخطاء.

# --hints--

يجب أن يحتوي استدعاء دالة `square` على معلمة `"something"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.result.value.matches("square('something')"));
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
