---
id: 699f3f636aa840b7201aa37c
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

إذا فتحت وحدة التحكم، سترى رسالة الخطأ التالية من TypeScript:

```md
Argument of type 'string' is not assignable to parameter of type 'number'
```

بما أن المعلمة مُعرفة صراحةً كرقم، يستطيع TypeScript اكتشاف الخطأ عند محاولة تمرير معلمة سلسلة نصية إلى الدالة `square`.

غيّر المعلمة الممررة إلى الدالة `square` إلى الرقم `5` لإصلاح الخطأ ورؤية الناتج الصحيح في وحدة التحكم.

# --hints--

يجب أن يحتوي استدعاء الدالة `square` على معلمة `5` بدلًا من `"something"`.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");
--fcc-editable-region--

console.log(result);
```

