---
id: 699f52220de88c8091032076
title: الخطوة 18
challengeType: 1
dashedName: step-18
---

# --description--

إذا فتحت وحدة التحكم، ستكون النتيجة `NaN` لأن المعلمة الثانية لـ `exponent` لم تُقدَّم أبدًا. هناك بعض التحديثات اللازمة لدالة `raiseTo` هذه لإصلاح ذلك.

التغيير الأول هو إرجاع `base ** 2`.

في الخطوة التالية، سترى كيفية التعامل مع المعلمة الاختيارية `exponent`.

# --hints--

يجب أن تُرجع دالة `raiseTo` رقمًا.

```js
assert.isNumber(raiseTo(4));
```

يجب أن تُرجع دالة `raiseTo` القيمة الناتجة عن رفع `base` إلى القوة `2`.

```js
assert.strictEqual(raiseTo(4), 16);
assert.strictEqual(raiseTo(5), 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
--fcc-editable-region--
  return base ** exponent;
--fcc-editable-region--
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
