---
id: 699f4639ae0016caaf81b7c6
title: الخطوة 10
challengeType: 1
dashedName: step-10
---

# --description--

حان الوقت الآن لاختبار دالتك.

ابدأ بإنشاء متغير اسمه `avgResult` وتعيين استدعاء الدالة `getAverage` له. يجب أن تكون المعلمة لهذا الاستدعاء هي `[2, 14, 26, 8]`.

ثم سجّل قيمة `avgResult` في وحدة التحكم.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن يكون لديك متغير اسمه `avgResult`.

```js
assert.isDefined(avgResult);
```

يجب أن تعيّن `getAverage([2, 14, 26, 8])` إلى متغير `avgResult` الخاص بك.

```js
assert.strictEqual(avgResult, getAverage([2, 14, 26, 8]));
```

يجب أن يكون متغير `avgResult` رقمًا.

```js
assert.isNumber(avgResult);
```

يجب أن تسجل `avgResult` في وحدة التحكم.

```js
assert.equal(getLogs()[1], avgResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
