---
id: 699f52fac083214b6bf51de6
title: الخطوة 19
challengeType: 1
dashedName: step-19
---

# --description--

في هذه الخطوة الأخيرة، تحتاج إلى إضافة تحقق شرطي لمعلمة `exponent` في الدالة `raiseTo`.

إذا تم توفير المعلمة `exponent`، يجب أن تُرجع الدالة `base` مرفوعة إلى قوة `exponent`. إذا لم تُوفر المعلمة `exponent`، يجب أن تُرجع `base ** 2`.

بهذه الخطوة الأخيرة، تكتمل مجموعة أدوات الرياضيات الآمنة من حيث النوع لديك!

# --hints--

يجب أن تُرجع دالة `raiseTo` القيمة `base ** exponent` عند توفير المعلمة `exponent`.

```js
assert.strictEqual(raiseTo(3, 3), 27);
```

يجب أن تُرجع دالة `raiseTo` القيمة `base ** 2` عند عدم توفير المعلمة `exponent`.

```js
assert.strictEqual(raiseTo(3), 9);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
--fcc-editable-region--
  
--fcc-editable-region--
  
  return base ** 2;
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  if (exponent) {
    return base ** exponent;
  }
  
  return base ** 2;
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
