---
id: 699f434e294a9e11cf4f65f0
title: الخطوة 8
challengeType: 1
dashedName: step-8
---

# --description--

حان الوقت الآن لإضافة مساعد رياضي آخر إلى مشروع أدواتك.

ابدأ بإنشاء دالة تسمى `getAverage` تأخذ معلمة واحدة. يجب أن تُرجع دالتك المتوسط لمصفوفة من الأرقام.

في الخطوات التالية، ستضيف تعليقات النوع لهذه الدالة.

# --hints--

يجب أن يكون لديك دالة تسمى `getAverage`.

```js
assert.isFunction(getAverage);
```

يجب أن تأخذ دالتك معلمة واحدة.

```js
assert.strictEqual(getAverage.length, 1);
```

يجب أن تُرجع دالتك رقمًا.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

يجب أن تُرجع دالتك النتيجة الصحيحة لأي مصفوفة من الأرقام.

```js
assert.strictEqual(getAverage([1, 2, 3, 4, 5]), 3);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
