---
id: 699f4ec939f10da62fb56c3e
title: الخطوة 15
challengeType: 1
dashedName: step-15
---

# --description--

حان الوقت الآن لاختبار دالتك.

أنشئ متغيرًا اسمه `raisedResult` وقم بتعيين نتيجة استدعاء `raiseTo` مع المعلمتين `2` و `3`.

ثم سجّل `raisedResult` في وحدة التحكم.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن يكون لديك متغير اسمه `raisedResult`.

```js
assert.isDefined(raisedResult);
```

يجب أن يكون متغير `raisedResult` رقمًا.

```js
assert.isNumber(raisedResult);
```

يجب أن تسجل `raisedResult` في وحدة التحكم.

```js
assert.equal(getLogs()[2], raisedResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
