---
id: 699f4861dd5b5f164e1bda4a
title: الخطوة 11
challengeType: 1
dashedName: step-11
---

# --description--

فائدة أخرى لاستخدام TypeScript هي ضمان تطبيق أنواع القيم المرجعة وأنها كما تتوقع. أحيانًا يستطيع TypeScript استنتاج نوع القيمة المرجعة من دالة، لكن يمكنك أيضًا تحديده صراحةً.

لتوضيح ذلك، غيّر نوع القيمة المرجعة من الدالة `getAverage` إلى السلسلة النصية `"something"`.

# --hints--

يجب أن تُرجع دالتك `getAverage` سلسلة نصية.

```js
assert.isString(getAverage([2, 14, 26, 8]));
```

يجب أن تُرجع دالتك `getAverage` السلسلة النصية `"something"`.

```js
assert.strictEqual(getAverage([2, 14, 26, 8]), "something");
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  --fcc-editable-region--
  return sum / numbers.length;
  --fcc-editable-region--
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
