---
id: 699f41b8da20f816b5c7dece
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

إذا فتحت وحدة التحكم، سترى رسالة الخطأ التالية:

```md
Expected 1 arguments, but got 2.
```

لحل الخطأ، حدّث استدعاء الدالة `square` ليأخذ معلمة واحدة فقط بقيمة `5`.

# --hints--

يجب أن يحتوي استدعاء الدالة `square` على معلمة واحدة فقط بقيمة `5`.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5, 10);
--fcc-editable-region--

console.log(result);
```

