---
id: 699f458e47fb88c756fdba81
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

في الوقت الحالي، معلمة `numbers` لها نوع ضمني `any`. لإصلاح ذلك، يمكنك تحديد نوع المعلمة صراحة كمصفوفة من الأرقام.

تذكر من الدروس السابقة، أنه يمكنك تحديد نوع مصفوفة من الأرقام بهذه الطريقة:

```ts
const numbers: number[] = [1, 2, 3];
```

حدّث دالة `getAverage` لتحدد نوع معلمة `numbers` صراحة كمصفوفة من الأرقام.

# --hints--

يجب أن تحتوي دالة `getAverage` على معلمة باسم `numbers` يتم تحديد نوعها صراحة كمصفوفة من الأرقام.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.getAverage.parameters;
assert.isTrue(parameters[0].matches("numbers: number[]"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers) {
--fcc-editable-region--
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
