---
id: 699f5063d2393c674b1f45dc
title: الخطوة 16
challengeType: 1
dashedName: step-16
---

# --description--

كما تتذكر من خطوة سابقة، إذا استدعيت دالة بعدد معلمات خاطئ، ستُرمى في TypeScript رسالة خطأ.

يمكنك جعل المعلمات اختيارية بإضافة `?` بعد اسم المعلمة. هذا يسمح لك باستدعاء الدالة بعدد معلمات أقل من المعلمات المعرفة.

إليك مثالاً:

```ts
function greet(name: string, greeting?: string) {
  if (greeting) {
    return `${greeting}, ${name}!`;
  }
  return `Hello, ${name}!`;
}
```

في هذا المثال، إذا لم تُزوّد معلمة `greeting` عند استدعاء الدالة `greet`، ستكون قيمتها `undefined`، وستُرجع الدالة تحية افتراضية.

حدّث دالتك `raiseTo` لجعل معلمة `exponent` اختيارية.

# --hints--

يجب أن تحتوي دالتك `raiseTo` على معلمة `exponent` معلمة كاختيارية.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[1].matches("exponent?: number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent: number): number {
--fcc-editable-region--
  return base ** exponent;
}

const raisedResult = raiseTo(2, 3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
