---
id: 699f4b903bf5427a599f791d
title: الخطوة 14
challengeType: 1
dashedName: step-14
---

# --description--

في الخطوات القليلة الماضية، ستبني دالة أخرى لمواصلة ممارسة توضيح أنواع البيانات مع الدوال.

ابدأ بإنشاء دالة تسمى `raiseTo` تأخذ معلمتين تسمى `base` و `exponent`. يجب أن تكون كلتا المعلمتين من نوع `number`. يجب أن تُرجع الدالة نتيجة رفع `base` إلى قوة `exponent`. يجب أيضًا توضيح نوع الإرجاع للدالة كـ `number`.

# --hints--

يجب أن يكون لديك دالة `raiseTo`.

```js
assert.isFunction(raiseTo);
```

يجب أن تحتوي دالة `raiseTo` على معلمتين، `base` و `exponent`.

```js
assert.match(
  code,
   /function\s+raiseTo\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*|(?:const|let)\s+raiseTo\s*=\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*(?::\s*\w+\s*)?\s*=>/
);
```

يجب أن تكون كل من المعلمتين `base` و `exponent` من نوع `number`.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[0].matches("base: number"));
assert.isTrue(parameters[1].matches("exponent: number"));
```

يجب أن تُرجع دالة `raiseTo` قيمة من نوع number.

```js
assert.isNumber(raiseTo(2, 3));
```

يجب أن تُرجع دالة `raiseTo` نتيجة رفع `base` إلى قوة `exponent`.

```js
assert.strictEqual(raiseTo(2, 3), 8);
```

يجب أن تحتوي دالة `raiseTo` على توضيح لنوع الإرجاع كـ `number`.

```js
const explorer = await __helpers.Explorer(code);
const { raiseTo } = explorer.allFunctions;
assert.isTrue(raiseTo.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
