---
id: 699f3cb8d057f29954205d7d
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

كما تتذكر من الدروس السابقة، يمكنك إضافة تعليقات نوع إلى معلمات الدالة بهذه الطريقة:

```ts
function add(num1: number, num2: number) {
  return num1 + num2;
}
```

في هذا المثال، سيعرض TypeScript خطأً إذا حاولت استدعاء الدالة `add` بمعلمات ليست أرقامًا.

حدّث دالتك `square` لتشمل تعليق نوع للمعلمة `num`. يجب أن يكون النوع `number`.

# --hints--

يجب إضافة تعليق نوع للمعلمة `num` في الدالة `square`. يجب أن يكون النوع `number`.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.square;
assert.isTrue(parameters[0].matches("num: number"));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
function square(num) {
--fcc-editable-region--
  return num * num;
}

const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
