---
id: 65c4f09e074dd8d37830ea00
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Unda darasa la vitu tupu lingine liitwalo `BinarySearchTree` ambalo linawakilisha mti wa utafutaji wa kibainari.

# --hints--

Unapaswa kuwa na darasa la vitu liitwalo `BinarySearchTree`. Kumbuka kutumia neno kuu la `pass`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(BinarySearchTree)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--
```
