---
id: 65c4f07da204f4d2f325fbd3
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Ndani ya njia ya `__init__`, anzisha sifa za kipengele za `left` na `right` za nodi kuwa `None`. Hii ni kwa sababu nodi inapoundwa kwa mara ya kwanza, haina watoto wa kushoto wala wa kulia. Kumbuka kutumia neno kuu la `self`.

# --hints--

Unapaswa kuweka thamani ya `None` kwa sifa ya kipengele ya `left` ya nodi kwa kutumia `self.left`.

```js
({ test: () => assert.match(code, /self\.left\s*=\s*None/m) })
```

Unapaswa kuweka thamani ya `None` kwa sifa ya kipengele ya `right` ya nodi kwa kutumia `self.right`.

```js
({ test: () => assert.match(code, /self\.right\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:    
    def __init__(self, key):
        self.key = key
--fcc-editable-region--
```
