---
id: 65c4f02dbd3557d23f12db62
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Ndani ya njia ya `__init__`, futa `pass` na weka thamani ya kigezo cha `key` kwenye sifa ya kipengele ya `key` ya nodi kwa kutumia `self.key`.

Hii ina maana kwamba sifa ya kipengele ya `key` ya mfano wa `TreeNode` itakuwa na thamani iliyopitishwa wakati wa kuunda kitu hicho.

# --hints--

Unapaswa kufuta tamko la `pass` kutoka kwenye njia ya `__init__`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("TreeNode").find_function("__init__").has_pass()`
      )
    );
  },
});

```

Unapaswa kuweka thamani ya kigezo cha `key` kwenye sifa ya kipengele ya `key` ya nodi kwa kutumia `self.key`.

```js
({ test: () => assert.match(code, /^\s{8}self\.key\s*=\s*key/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__(self, key):
        pass
--fcc-editable-region--
```
