---
id: 65c4f0d32ca17ad4b1636b0e
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Ndani ya njia ya `__init__`, futa `pass` na anzisha `root` kwa thamani ya `None`.

Sifa ya kipengele `root` inawakilisha nodi mzizi wa mti wa utafutaji wa kibainari. Kwa kuwa huu ni kijenzi wakati kitu kipya cha `BinarySearchTree` kinapotengenezwa, huanza na mti tupu, hivyo sifa ya kipengele `root` imesanidiwa kuwa `None`.

# --hints--

Unapaswa kuondoa tamko la `pass` kutoka kwenye njia ya `__init__`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

Unapaswa kuanzisha sifa ya kipengele `root` kuwa `None` kwa kutumia `self.root`.

```js
({ test: () => assert.match(code, /^\s{8}self\.root\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        pass
--fcc-editable-region--        
```
