---
id: 65c4f37ba6cbcfdc77e24165
title: Hatua 17
challengeType: 20
dashedName: step-17
---

# --description--

Njia ya `insert` itaitwa na mtumizi. Mbali na kigezo cha `self`, pia itahitaji kigezo cha `key`. Kigezo hiki kitakuwa thamani ya ufunguo ya kuingiza katika mti wa utafutaji wa kibainari.

Ongeza `key` kama kigezo cha pili katika ufafanuzi wa kitendakazi.

# --hints--

Njia ya `insert` inapaswa kuwa na kigezo cha `self`.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,/) });
```

Njia ya `insert` inapaswa kuwa na kigezo cha `key`.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

--fcc-editable-region--            
    def insert(self):
        pass

--fcc-editable-region--

```
