---
id: 65c4f0b82db314d3fcc0b8e1
title: Hatua ya 7
challengeType: 20
dashedName: step-7
---

# --description--

Ndani ya darasa la `BinarySearchTree`, badilisha `pass` kwa njia ya `__init__` na ongeza kigezo cha `self` kwa njia hii.

# --hints--

Unapaswa kuondoa neno la `pass` kutoka kwenye darasa la `BinarySearchTree`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});

```


Unapaswa kufafanua njia ya `__init__` na kuongeza kigezo cha `self` kwa njia hii. Kumbuka kutumia neno la `pass` ndani ya mwili wa njia.

```js
({ test: () => assert.match(code, /^\s{4}def\s+__init__\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
        
--fcc-editable-region--
class BinarySearchTree:
    pass
--fcc-editable-region--
```
