---
id: 65c4f195de7e2ad5932be717
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

Ifuatayo, unahitaji kufafanua mbinu ya kuingiza nodi kwenye mti. Kwa ajili ya hilo, unahitaji kufafanua njia ya `_insert`, ambayo ni kitendakazi cha msaada na itatumika na njia halisi ya `insert` baadaye.

Njia hii ni ya kurudiwa, ikimaanisha inajipigia wito yenyewe ili kupita kwenye mti hadi mahali panapofaa pa nodi mpya kupatikana.

Fafanua njia ya `_insert` yenye vigezo `self`, `node` na `key`.

# --hints--

Unapaswa kufafanua njia ya `_insert` ndani ya darasa la vitu la `BinarySearchTree`. Kumbuka kutumia `pass`.

```js
({ test: () => assert.match(code, /def\s+_insert\s*\([^(]*\)\s*:/m) })
```

Njia yako ya `_insert` inapaswa kuchukua vigezo vitatu: `self`, `node` na `key`.

```js
({ test: () => assert.match(code, /def\s+_insert\s*\(\s*self\s*,\s*node\s*,\s*key\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--
```
