---
id: 65c4f2851645e8d84d14f5d1
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Sasa unahitaji kupita mti wa kibainari kwa kurudia na kuingiza thamani kwa kutumia kanuni za miti ya kibainari:

- Thamani ndogo kuliko funguo huwekwa katika mti mdogo wa kushoto
- Thamani kubwa kuliko funguo huwekwa katika mti mdogo wa kulia

Baada ya taarifa yako ya if iliyopo, andika taarifa nyingine ya `if` ili kukagua kama `key` ni ndogo kuliko `node.key`.

# --hints--

Unapaswa kuandika taarifa nyingine ya if ili kukagua kama `key` ni ndogo kuliko `node.key`. Kumbuka kutumia neno kuu la `pass`.

```js
({ test: () => assert(runPython(` _Node(_code).find_class("BinarySearchTree").find_function("_insert").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

--fcc-editable-region--
```
