---
id: 65c4ef8463f869d0eaf87c67
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Ndani ya darasa la `TreeNode`, badilisha `pass` na njia ya `__init__` ili uweze kuanzisha sifa za kipengele za kitu. Usiongeze vigezo kwa sasa.

# --hints--

Unapaswa kuondoa neno kuu la `pass` kutoka darasa la `TreeNode` na kulihamisha ndani ya njia ya `__init__`.

```js
({
  test: () => {
    const pyClassStr = runPython(`str(_Node(_code).find_class("TreeNode"))`);
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});
```

Unapaswa kufafanua njia ya `__init__` ndani ya darasa lako la `TreeNode`. Kumbuka kutumia neno kuu la `pass` ndani yake.

```js
({
  test: () => {
    assert(
      runPython(`_Node(_code).find_class("TreeNode").has_function("__init__")`)
    );
  },
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
--fcc-editable-region--
```
