---
id: 65c4eb814cc977c95cd7df0e
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika mradi huu, utaunda mti wa utafutaji wa kibainari (BST). BST ni muundo wa data ambapo kila nodi ina watoto wawili angalau, mtoto wa kushoto akiwa na thamani ndogo kuliko nodi ya mzazi na mtoto wa kulia akiwa na thamani kubwa kuliko nodi ya mzazi, jambo linaloruhusu utendaji mzuri wa utafutaji na upangaji.

Huu ndio muonekano wa mti wa utafutaji wa kibainari:

<img alt="a binary search tree" src="https://cdn.freecodecamp.org/curriculum/python/bst-example.png" style="background-color: white; height:500px; width:500px; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Anza kwa kutangaza darasa la vitu tupu `TreeNode`. Darasa la vitu `TreeNode` linawakilisha nodi katika mti wa utafutaji wa kibainari. Tumia neno kuu `pass` kujaza mwili wa darasa na kuepuka kosa.

# --hints--

Unapaswa kutumia neno kuu `class` kutangaza darasa la vitu tupu lenye jina `TreeNode`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(TreeNode)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
