---
id: 65c4f2b7178afed8e88f782f
title: Hatua 12
challengeType: 20
dashedName: step-12
---

# --description--

Ikiwa `key` < `node.key` itathibitishwa kuwa `True`, basi nodi mpya inapaswa kuwekwa katika mti mdogo wa kushoto.

Futa `pass` na fanya wito wa njia ya `_insert` kwa kurudia ukitumia mtoto wa kushoto kama hoja ya kwanza na `key` kama hoja ya pili. Weka thamani ya matokeo kwenye sifa ya kipengele ya `left` ya nodi ya sasa.

# --hints--

Unapaswa kufuta neno kuu la `pass` kutoka kwenye blokki ya `if`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_insert").find_if("key < node.key").find_bodies()[0].has_pass()`
      )
    );
  },
});
```

Unapaswa kuita njia ya `self._insert` ukipitisha `node.left` na `key` kama hoja.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

Unapaswa kuweka thamani ya matokeo ya wito lako la `self._insert()` kwenye sifa ya kipengele ya `left` ya nodi ya sasa.

```js
({ test: () => assert.match(code, /node\.left\s*=\s*self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--    

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            pass
--fcc-editable-region--    

```
