---
id: 65c4f3258d2e4cdacc919dfd
title: Hatua 15
challengeType: 20
dashedName: step-15
---

# --description--

Mwisho wa njia yako ya `_insert`, baada ya mchakato wa kuingiza kukamilika, rudisha nodi ya sasa ili kusasisha muundo wa mti katika ngazi za juu za orodha ya wito za mzunguko.

# --hints--

Unapaswa kurudisha nodi ya sasa nje ya vifungu vya masharti.

```js
({ test: () => assert.match(code, /^\s{8}return\s+node/m) })
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--    
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)

--fcc-editable-region--

```
