---
id: 65c4f013851cefd1a4fe4c96
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Njia ya `__init__` inachukua vigezo viwili: `self` (ambayo inawakilisha mfano wa darasa la vitu unaoundwa) na `key` (thamani inayohifadhiwa katika nodi). Ongeza vigezo hivyo viwili kwenye njia ya `__init__()`.

# --hints--

Unapaswa kuongeza kigezo cha `self` kwenye njia.

```js
({ test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*/) });
```

Unapaswa kuongeza kigezo cha `key` kwenye njia.

```js
({
  test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__():
        pass
--fcc-editable-region--
```
