---
id: 65c6409418806da8d0636ffc
title: Hatua 26
challengeType: 20
dashedName: step-26
---

# --description--

Kazi za `insert` na `search` zimekamilika, ni wakati wa kuzijaribu.

Kumbuka kwamba, kwa sasa, nodi hazijasawazishwa na zimeingizwa tu. Utafanya kazi ya kusawazisha kwa kutumia utambazaji wa inorder baadaye.

Unaweza kuunda mfano wa darasa la vitu katika Python kama ifuatavyo:

```python
object_name = ClassName()
```

Nje ya ufafanuzi wa madarasa, unda mfano wa darasa la `BinarySearchTree` na uweke thamani yake katika mabadiliko `bst`.


# --hints--

Unapaswa kuunda mfano wa darasa la `BinarySearchTree`.

```js
assert.match(code, /BinarySearchTree\(\s*\)/);
```

Unapaswa kuweka thamani mfano mpya wa `BinarySearchTree` katika mabadiliko `bst`.

```js
assert.match(code, /bst\s*=\s*BinarySearchTree\(\s*\)/);
``` 

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--


--fcc-editable-region--
```
