---
id: 65c63a4c4da62e9ae18e321a
title: Hatua 19
challengeType: 20
dashedName: step-19
---

# --description--

Ni wakati wa kufanya kazi kwenye utendaji wa utafutaji. Kama ulivyotengeneza njia ya msaada `_insert` kwa njia ya `insert`, unahitaji kutengeneza njia ya msaada `_search` kwa njia ya `search`.

Tangaza njia ya `_search` yenye vigezo vitatu, yaani `self`, `node` na `key`.

# --hints--

Unapaswa kutangaza njia ya `_search`. Kumbuka mstari wa mstari wa chini na kumbuka kutumia `pass`.

```js
({ test: () => assert.match(code, /def\s+_search/) });
```

Njia yako ya `_search` inapaswa kuchukua vigezo vitatu: `self`, `node`, na `key`.

```js
({ test: () => assert.match(code, /def\s+_search\(\s*self\s*,\s*node\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
--fcc-editable-region--

--fcc-editable-region--
```
