---
id: 65c63da2ed6769a10e141341
title: Hatua 20
challengeType: 20
dashedName: step-20
---

# --description--

Sasa utafafanua kesi ya msingi kwa ajili ya utafutaji wa kurudia. Ondoa `pass` iliyopo na andika usemi wa `if` unaochunguza masharti mawili:

- Ikiwa `node` ni `None`: Hii inaonyesha kuwa utafutaji umefikia mwisho wa tawi bila kupata funguo.
- Ikiwa `node.key == key`: Hii inamaanisha kuwa funguo imepatikana katika nodi ya sasa.

Unganisha masharti hayo mawili kwa kutumia kielezi cha `or` na rudisha nodi ya sasa ndani ya blokki ya `if`.

# --hints--

Unapaswa kuondoa neno kuu la `pass` kutoka kwenye njia ya `_search`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_search").has_pass()`
      )
    );
  },
});

```

Unapaswa kuandika usemi wa `if` unaochunguza kama `node` ni `None` au kama `node.key` ni sawa na `key`.

```js
({ test: () => assert.match(code, /^\s{8}if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:/m) })
```

Unapaswa kurudisha `node` kutoka kwenye blokki ya `if`.

```js
({ test: () => assert.match(code, /^(\s{8})if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:\s*^\1\s{4}return\s+node/m) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        pass

--fcc-editable-region--
```
