---
id: 65c63fd9b387f0a6c5bf6a72
title: Hatua 24
challengeType: 20
dashedName: step-24
---

# --description--

Ndani ya njia ya `search`, futa `pass` na fanya wito kwa njia ya msaidizi `_search` kwa hoja zifuatazo.

- `self.root`: Hii ni mzizi wa mti wa utafutaji wa kibainari. Utafutaji huanza kutoka mzizi.
- `key`: Hii ni thamani ambayo mtumizi anataka kuipata katika mti wa utafutaji wa kibainari.

Ndani, `search` huwasilisha mantiki halisi ya utafutaji kwa njia ya msaidizi `_search` inayofanya utafutaji wa kurudia ndani ya mti wa utafutaji wa kibainari.

# --hints--

Unapaswa kuondoa neno la `pass` kutoka njia ya `search`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("search").has_pass()`
      )
    );
  },
});

```

Unapaswa kuita njia ya `_search` ndani ya njia ya `search`.

```js
({ test: () => assert(runPython(`
  str(_Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0]).startswith("self._search(")
`)) });
```

Unapaswa kuita njia ya `_search` ukipitisha `self.root` na `key` kama hoja.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0].is_equivalent("self._search(self.root, key)")
`)) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        pass
--fcc-editable-region--
```
