---
id: 65ca07dd6fa8840491b7a5cd
title: Hatua 41
challengeType: 20
dashedName: step-41
---

# --description--

Kama hakuna mojawapo ya masharti yaliyotangulia yanayotekelezwa, ina maana nodi ina watoto wa kushoto na kulia.

Ili kuchagua mrithi, unahitaji kupata thamani ya kiwango cha chini katika mti mdogo wa kulia.
Thamani ndogo kabisa itakuwa mrithi kwa mpangilio wa nodi ya sasa.

Ili kupata thamani ndogo kabisa, tengeneza kitendakazi cha msaada `_min_value` kinachopokea vigezo viwili: `self` na `node`.

# --hints--

Unapaswa kufafanua njia `_min_value` ukiwa na `self` na `node` kama vigezo. Kumbuka kutumia neno kuu la `pass`.

```js
({ test: () => assert.match(code, /def\s+_min_value\(\s*self\s*,\s*node\s*\)\s*:/) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left   
--fcc-editable-region--
    
--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)
    
# print('Search for 80:', bst.search(80))


```
