---
id: 65d8a6fcb15a3a239ba35dfd
title: Hatua 46
challengeType: 20
dashedName: step-46
---

# --description--

Baada ya kupata kiwango cha chini, utahitaji kufuta nodi yenye kiwango cha chini kutoka kwenye mti mdogo wa kulia kwa njia ya kurudia.

Hatua hii inahakikisha kuwa nodi yenye kiwango cha chini inafutwa kutoka kwenye mti huku ikidumisha sifa ya mti wa utafutaji wa kibainari (BST).

Fanya wito wa njia ya `_delete` kwa kurudia ukitumia `node.right` na `node.key` kama hoja. Weka thamani ya return ya wito wa `_delete()` kwa mtoto wa kulia wa nodi ya sasa.

# --hints--

Unapaswa kufanya wito wa njia ya `_delete` kwa kurudia ukitumia `node.right` na `node.key` kama hoja.

```js
assert.match(code, /self\._delete\(\s*node\.right\s*,\s*node\.key\s*/);
```

Unapaswa kuweka thamani ya return ya wito wa `_delete()` kwa mtoto wa kulia wa nodi ya sasa.

```js
assert.match(code, /node\.right\s*=\s*self\._delete\(\s*node\.right\s*,\s*node\.key/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key)
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left

            node.key = self._min_value(node.right) 
            
--fcc-editable-region--
    def _min_value(self, node):
        while node.left is not None:
            node = node.left

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
