---
id: 587d7dbd367417b2b2512bb6
title: Wiederverwendbares CSS mit Mixins erstellen
challengeType: 0
forumTopicId: 301455
dashedName: create-reusable-css-with-mixins
---

# --description--

In Sass ist ein <dfn>Mixin</dfn> eine Gruppe von CSS-Deklarationen, die im gesamten Stylesheet wiederverwendet werden können. Die Definition beginnt mit der `@mixin`-At-Regel, gefolgt von einem benutzerdefinierten Namen. Sie wenden das Mixin mit der `@include`-At-Regel an.

```scss
@mixin reset-list {
  margin: 0;
  padding: 0;
  list-style: none;
}

nav ul {
  @include reset-list;
}
```

Kompiliert zu:

```css
nav ul {
  margin: 0;
  padding: 0;
  list-style: none;
}
```

Ihre Mixins können auch Argumente annehmen, wodurch ihr Verhalten angepasst werden kann. Die Argumente sind beim Verwenden des Mixins erforderlich.

```scss
@mixin prose($font-size, $spacing) {
  font-size: $font-size;
  margin: 0;
  margin-block-end: $spacing;
}

p {
  @include prose(1.25rem, 1rem);
}

h2 {
  @include prose(2.4rem, 1.5rem);
}
```

Sie können Argumente optional machen, indem Sie den Parametern Standardwerte zuweisen.

```scss
@mixin text-color($color: black) {
  color: $color;
}

p {
  @include text-color(); /* color: black */
}

nav a {
  @include text-color(orange);
}
```

# --instructions--

Schreiben Sie ein Mixin mit dem Namen `shape` und geben Sie ihm 3 Parameter: `$w`, `$h` und `$bg-color`.

Verwenden Sie das `shape`-Mixin, um dem `#square`-Element eine Breite und Höhe von `50px` sowie die Hintergrundfarbe `red` zu geben. Für das `#rect-a`-Element fügen Sie eine Breite von `100px`, eine Höhe von `50px` und die Hintergrundfarbe `blue` hinzu. Schließlich fügen Sie für das `#rect-b`-Element eine Breite von `50px`, eine Höhe von `100px` und die Hintergrundfarbe `orange` hinzu.

# --hints--

Sie sollten ein Mixin mit dem Namen `shape` mit 3 Parametern deklarieren: `$w`, `$h` und `$bg-color`.

```js
assert.match(code, /@mixin\s+shape\s*\(\s*\$w,\s*\$h,\s*\$bg-color\s*\)\s*{/gi);
```

Ihr Mixin sollte eine `width`-Eigenschaft enthalten, die den `$w`-Parameter verwendet.

```js
assert.match(__helpers.removeWhiteSpace(code), /width:\$w;/gi);
```

Ihr Mixin sollte eine `height`-Eigenschaft enthalten, die den `$h`-Parameter verwendet.

```js
assert.match(__helpers.removeWhiteSpace(code), /height:\$h;/gi);
```

Ihr Mixin sollte eine `background-color`-Eigenschaft enthalten, die den `$bg-color`-Parameter verwendet.

```js
assert.match(__helpers.removeWhiteSpace(code), /background-color:\$bg\-color;/gi);
```

Sie sollten die Styles innerhalb des `#square`-Selektors durch einen Aufruf des `shape`-Mixins mit dem Schlüsselwort `@include` ersetzen. Setzen Sie eine Breite und Höhe von `50px` sowie die Hintergrundfarbe `red`.

```js
assert.match(code, /#square\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*50px\s*,\s*red\s*\)\s*;\s*}/gi);
```

Sie sollten die Styles innerhalb des `#rect-a`-Selektors durch einen Aufruf des `shape`-Mixins mit dem Schlüsselwort `@include` ersetzen. Setzen Sie eine Breite von `100px`, eine Höhe von `50px` und die Hintergrundfarbe `blue`.

```js
assert.match(code, /#rect-a\s*{\s*@include\s+shape\s*\(\s*100px\s*,\s*50px\s*,\s*blue\s*\)\s*;\s*}/gi);
```

Sie sollten die Styles innerhalb des `#rect-b`-Selektors durch einen Aufruf des `shape`-Mixins mit dem Schlüsselwort `@include` ersetzen. Setzen Sie eine Breite von `50px`, eine Höhe von `100px` und die Hintergrundfarbe `orange`.

```js
assert.match(code, /#rect-b\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*100px\s*,\s*orange\s*\)\s*;\s*}/gi);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
#square {
  width: 50px;
  height: 50px;
  background-color: red;
}

#rect-a {
  width: 100px;
  height: 50px;
  background-color: blue;
}

#rect-b {
  width: 50px;
  height: 100px;
  background-color: orange;
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```

# --solutions--

```html
<style type='text/scss'>
@mixin shape($w, $h, $bg-color) {
  width: $w;
  height: $h;
  background-color: $bg-color;
}

#square {
  @include shape(50px, 50px, red);
}

#rect-a {
  @include shape(100px, 50px, blue);
}

#rect-b {
  @include shape(50px, 100px, orange);
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```
