---
id: 587d7dbe367417b2b2512bb9
title: "Verwenden Sie @for, um eine Sass-Schleife zu erstellen"
challengeType: 0
forumTopicId: 301462
dashedName: use-for-to-create-a-sass-loop
---

# --description--

Die `@for`-Direktive fügt Stile in einer Schleife hinzu, sehr ähnlich wie eine `for`-Schleife in JavaScript.

`@for` wird auf zwei Arten verwendet: „start bis end“ oder „start durch end“. Der Hauptunterschied besteht darin, dass „start **bis** end“ die Endzahl *ausschließt* und „start **durch** end“ die Endzahl *einschließt*.

Hier ist ein Beispiel für „start **durch** end“:

```scss
@for $i from 1 through 12 {
  .col-#{$i} { width: 100%/12 * $i; }
}
```

Der `#{$i}`-Teil ist die Syntax, um eine Variable (`i`) mit Text zu kombinieren und so einen String zu erzeugen. Wenn die Sass-Datei in CSS umgewandelt wird, sieht das so aus:

```scss
.col-1 {
  width: 8.33333%;
}

.col-2 {
  width: 16.66667%;
}

...

.col-12 {
  width: 100%;
}
```

Dies ist eine leistungsstarke Methode, um ein Rasterlayout zu erstellen. Nun haben Sie zwölf Optionen für Spaltenbreiten als CSS-Klassen zur Verfügung.

# --instructions--

Schreiben Sie eine `@for`-Direktive, die eine Variable `$j` verwendet, die von 1 **bis** 6 läuft.

Sie soll 5 Klassen namens `.text-1` bis `.text-5` erzeugen, wobei jede eine `font-size` hat, die auf 15px multipliziert mit dem Index gesetzt ist.

# --hints--

Ihr Code sollte die `@for`-Direktive verwenden.

```js
assert(code.match(/@for /g));
```

Ihre `.text-1`-Klasse sollte eine `font-size` von 15px haben.

```js
assert($('.text-1').css('font-size') == '15px');
```

Ihre `.text-2`-Klasse sollte eine `font-size` von 30px haben.

```js
assert($('.text-2').css('font-size') == '30px');
```

Ihre `.text-3`-Klasse sollte eine `font-size` von 45px haben.

```js
assert($('.text-3').css('font-size') == '45px');
```

Ihre `.text-4`-Klasse sollte eine `font-size` von 60px haben.

```js
assert($('.text-4').css('font-size') == '60px');
```

Ihre `.text-5`-Klasse sollte eine `font-size` von 75px haben.

```js
assert($('.text-5').css('font-size') == '75px');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

# --solutions--

```html
<style type='text/scss'>

@for $i from 1 through 5 {
  .text-#{$i} { font-size: 15px * $i; }
}

</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

---

```html
<style type='text/scss'>

@for $i from 1 to 6 {
  .text-#{$i} { font-size: 15px * $i; }
}

</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```
