---
id: 587d7dbd367417b2b2512bb4
title: Daten mit Sass-Variablen speichern
challengeType: 0
forumTopicId: 301460
dashedName: store-data-with-sass-variables
---

# --description--

Eine Funktion von Sass, die sich von CSS unterscheidet, ist die Verwendung von Variablen. Diese werden deklariert und zum Speichern von Daten verwendet, ähnlich wie in JavaScript.

In JavaScript werden Variablen mit den Schlüsselwörtern `let` und `const` definiert. In Sass beginnen Variablen mit einem `$`, gefolgt vom Variablennamen.

Hier sind ein paar Beispiele:

```scss
$main-fonts: Arial, sans-serif;
$headings-color: green;
```

Und so verwendet man die Variablen:

```scss
h1 {
  font-family: $main-fonts;
  color: $headings-color;
}
```

Ein Beispiel, bei dem Variablen nützlich sind, ist, wenn mehrere Elemente dieselbe Farbe haben sollen. Wenn sich diese Farbe ändert, muss nur der Wert der Variablen angepasst werden.

# --instructions--

Erstellen Sie eine Variable `$text-color` und setzen Sie sie auf `red`. Ändern Sie dann den Wert der `color`-Eigenschaft für `.blog-post` und `h2` auf die `$text-color`-Variable.

# --hints--

Ihr Code sollte eine Sass-Variable für `$text-color` mit dem Wert `red` deklariert haben.

```js
assert(code.match(/\$text-color\s*:\s*?red\s*;/g));
```

Ihr Code sollte die `$text-color`-Variable verwenden, um die `color` für die `.blog-post` und `h2` Elemente zu ändern.

```js
assert(code.match(/color\s*:\s*\$text-color\s*;?/g));
```

Ihr `.blog-post`-Element sollte eine `color` von rot haben.

```js
assert($('.blog-post').css('color') == 'rgb(255, 0, 0)');
```

Ihre `h2`-Elemente sollten eine `color` von rot haben.

```js
assert($('h2').css('color') == 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>


  .header{
    text-align: center;
  }
  .blog-post, h2 {
    color: red;
  }
</style>

<h1 class="header">Learn Sass</h1>
<div class="blog-post">
  <h2>Some random title</h2>
  <p>This is a paragraph with some random text in it</p>
</div>
<div class="blog-post">
  <h2>Header #2</h2>
  <p>Here is some more random text.</p>
</div>
<div class="blog-post">
  <h2>Here is another header</h2>
  <p>Even more random text within a paragraph</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  $text-color: red;

  .header{
    text-align: center;
  }
  .blog-post, h2 {
    color: $text-color;
  }
</style>

<h1 class="header">Learn Sass</h1>
<div class="blog-post">
  <h2>Some random title</h2>
  <p>This is a paragraph with some random text in it</p>
</div>
<div class="blog-post">
  <h2>Header #2</h2>
  <p>Here is some more random text.</p>
</div>
<div class="blog-post">
  <h2>Here is another header</h2>
  <p>Even more random text within a paragraph</p>
</div>
```
