---
id: 587d7dbd367417b2b2512bb5
title: CSS mit Sass verschachteln
challengeType: 0
forumTopicId: 301457
dashedName: nest-css-with-sass
---

# --description--

Sass erlaubt das Verschachteln von CSS-Regeln, was eine nützliche Technik zur Organisation eines Stylesheets ist.

Normalerweise wird jedes Element in einer eigenen Zeile angesprochen, um es zu stylen, so:

```scss
article {
  height: 200px;
}

article p {
  color: white;
}

article ul {
  color: blue;
}
```

Bei einem großen Projekt enthält die CSS-Datei viele Zeilen und Regeln. Hier kann das Verschachteln helfen, Ihren Code zu organisieren, indem Kind-Stilregeln innerhalb der jeweiligen Elternelemente platziert werden:

```scss
article {
  height: 200px;

  p {
    color: white;
  }

  ul {
    color: blue;
  }
}

```

# --instructions--

Verwenden Sie die oben gezeigte Verschachtelungstechnik, um die CSS-Regeln für beide Kinder des `.blog-post`-Elements neu zu organisieren. Zum Testen sollte das `h1`-Element vor dem `p`-Element stehen.

# --hints--

Ihr Code sollte die CSS-Regeln so neu organisieren, dass `h1` und `p` im Elternelement `.blog-post` verschachtelt sind.

```js
assert(
  code.match(
    /\.blog-post\s*?{\s*?h1\s*?{\s*?text-align:\s*?center;\s*?color:\s*?blue;\s*?}\s*?p\s*?{\s*?font-size:\s*?20px;\s*?}\s*?}/gi
  )
);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
  .blog-post {

  }
  h1 {
    text-align: center;
    color: blue;
  }
  p {
    font-size: 20px;
  }
</style>

<div class="blog-post">
  <h1>Blog Title</h1>
  <p>This is a paragraph</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  .blog-post {
    h1 {
      text-align: center;
      color: blue;
    }
    p {
      font-size: 20px;
    }
  }
</style>

<div class="blog-post">
  <h1>Blog Title</h1>
  <p>This is a paragraph</p>
</div>
```
