---
id: 587d7fa5367417b2b2512bbd
title: Eine CSS-Stilmenge auf ein anderes Element erweitern
challengeType: 0
forumTopicId: 301456
dashedName: extend-one-set-of-css-styles-to-another-element
---

# --description--

Sass hat eine Funktion namens `extend`, die es einfach macht, die CSS-Regeln eines Elements zu übernehmen und für ein anderes weiterzuverwenden.

Zum Beispiel stylt der folgende Block von CSS-Regeln eine `.panel`-Klasse. Sie hat eine `background-color`, `height` und `border`.

```scss
.panel{
  background-color: red;
  height: 70px;
  border: 2px solid green;
}
```

Nun möchten Sie ein weiteres Panel namens `.big-panel`. Es hat dieselben Basiseigenschaften wie `.panel`, benötigt aber auch eine `width` und `font-size`. Es ist möglich, die anfänglichen CSS-Regeln von `.panel` zu kopieren und einzufügen, aber der Code wird wiederholend, wenn Sie mehr Paneltypen hinzufügen. Die `extend`-Direktive ist eine einfache Möglichkeit, die für ein Element geschriebenen Regeln wiederzuverwenden und dann weitere für ein anderes hinzuzufügen:

```scss
.big-panel{
  @extend .panel;
  width: 150px;
  font-size: 2em;
}
```

Die `.big-panel` wird dieselben Eigenschaften wie `.panel` haben, zusätzlich zu den neuen Stilen.

# --instructions--

Erstellen Sie eine Klasse `.info-important`, die `.info` erweitert und außerdem eine `background-color` mit dem Wert magenta hat.

# --hints--

Ihre `info-important`-Klasse sollte eine `background-color` mit dem Wert `magenta` haben.

```js
assert(
  code.match(
    /\.info-important\s*?{[\s\S]*background-color\s*?:\s*?magenta\s*?;[\s\S]*}/gi
  )
);
```

Ihre `info-important`-Klasse sollte `@extend` verwenden, um das Styling von der `info`-Klasse zu erben.

```js
assert(
  code.match(/\.info-important\s*?{[\s\S]*@extend\s*?.info\s*?;[\s\S]*/gi)
);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
  h3{
    text-align: center;
  }
  .info{
    width: 200px;
    border: 1px solid black;
    margin: 0 auto;
  }




</style>
<h3>Posts</h3>
<div class="info-important">
  <p>This is an important post. It should extend the class ".info" and have its own CSS styles.</p>
</div>

<div class="info">
  <p>This is a simple post. It has basic styling and can be extended for other uses.</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  h3{
    text-align: center;
  }
  .info{
    width: 200px;
    border: 1px solid black;
    margin: 0 auto;
  }
  .info-important{
    @extend .info;
    background-color: magenta;
  }



</style>
<h3>Posts</h3>
<div class="info-important">
  <p>This is an important post. It should extend the class ".info" and have its own CSS styles.</p>
</div>

<div class="info">
  <p>This is a simple post. It has basic styling and can be extended for other uses.</p>
</div>
```
