---
id: 587d7dbe367417b2b2512bb8
title: "Verwenden Sie @if und @else, um Logik zu Ihren Styles hinzuzufügen"
challengeType: 0
forumTopicId: 301463
dashedName: use-if-and-else-to-add-logic-to-your-styles
---

# --description--

Die `@if`-Direktive in Sass ist nützlich, um einen bestimmten Fall zu testen – sie funktioniert genau wie die `if`-Anweisung in JavaScript.

```scss
@mixin make-bold($bool) {
  @if $bool == true {
    font-weight: bold;
  }
}
```

Und genau wie in JavaScript testen die `@else if`- und `@else`-Direktiven weitere Bedingungen:

```scss
@mixin text-effect($val) {
  @if $val == danger {
    color: red;
  }
  @else if $val == alert {
    color: yellow;
  }
  @else if $val == success {
    color: green;
  }
  @else {
    color: black;
  }
}
```

# --instructions--

Erstellen Sie ein Mixin namens `border-stroke`, das einen Parameter `$val` entgegennimmt. Das Mixin soll mit den `@if`-, `@else if`- und `@else`-Direktiven die folgenden Bedingungen prüfen:

```scss
light - 1px solid black
medium - 3px solid black
heavy - 6px solid black
```

Wenn der Wert des Parameters `$val` nicht `light`, `medium` oder `heavy` ist, soll die `border`-Eigenschaft auf `none` gesetzt werden.

# --hints--

Ihr Code soll ein Mixin namens `border-stroke` deklarieren, das einen Parameter namens `$val` hat.

```js
assert(code.match(/@mixin\s+?border-stroke\s*?\(\s*?\$val\s*?\)\s*?{/gi));
```

Ihr Mixin soll eine `@if`-Anweisung enthalten, die prüft, ob `$val` `light` ist, und die `border` auf `1px solid black` setzt.

```js
assert(
  code.match(
    /@if\s+?\$val\s*?===?\s*?light\s*?{\s*?border\s*?:\s*?1px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

Ihr Mixin soll eine `@else if`-Anweisung enthalten, die prüft, ob `$val` `medium` ist, und die `border` auf `3px solid black` setzt.

```js
assert(
  code.match(
    /@else\s+?if\s+?\$val\s*?===?\s*?medium\s*?{\s*?border\s*?:\s*?3px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

Ihr Mixin soll eine `@else if`-Anweisung enthalten, die prüft, ob `$val` `heavy` ist, und die `border` auf `6px solid black` setzt.

```js
assert(
  code.match(
    /@else\s+?if\s+?\$val\s*?===?\s*?heavy\s*?{\s*?border\s*?:\s*?6px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

Ihr Mixin soll eine `@else`-Anweisung enthalten, die die `border` auf `none` setzt.

```js
assert(code.match(/@else\s*?{\s*?border\s*?:\s*?none\s*?;\s*?}/gi));
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



  #box {
    width: 150px;
    height: 150px;
    background-color: red;
    @include border-stroke(medium);
  }
</style>

<div id="box"></div>
```

# --solutions--

```html
<style type='text/scss'>
  @mixin border-stroke($val) {
    @if $val == light {
      border: 1px solid black;
    }
    @else if $val == medium {
      border: 3px solid black;
    }
    @else if $val == heavy {
      border: 6px solid black;
    }
    @else {
      border: none;
    }
  }


  #box {
    width: 150px;
    height: 150px;
    background-color: red;
    @include border-stroke(medium);
  }
</style>

<div id="box"></div>
```
