---
id: 587d7dbf367417b2b2512bba
title: "Verwenden Sie @each, um Elemente in einer Liste zu durchlaufen"
challengeType: 0
forumTopicId: 301461
dashedName: use-each-to-map-over-items-in-a-list
---

# --description--

Die letzte Aufgabe zeigte, wie die `@for`-Direktive einen Start- und Endwert verwendet, um eine bestimmte Anzahl von Wiederholungen durchzuführen. Sass bietet auch die `@each`-Direktive, die über jedes Element in einer Liste oder Map iteriert. Bei jeder Iteration wird der Variablen der aktuelle Wert aus der Liste oder Map zugewiesen.

```scss
@each $color in blue, red, green {
  .#{$color}-text {color: $color;}
}
```

Eine Map hat eine etwas andere Syntax. Hier ist ein Beispiel:

```scss
$colors: (color1: blue, color2: red, color3: green);

@each $key, $color in $colors {
  .#{$color}-text {color: $color;}
}
```

Beachten Sie, dass die `$key`-Variable benötigt wird, um auf die Schlüssel in der Map zu verweisen. Andernfalls würde der kompilierte CSS `color1`, `color2` ... enthalten. Beide obigen Codebeispiele werden in den folgenden CSS-Code umgewandelt:

```scss
.blue-text {
  color: blue;
}

.red-text {
  color: red;
}

.green-text {
  color: green;
}
```

# --instructions--

Schreiben Sie eine `@each`-Direktive, die eine Liste durchläuft: `blue, black, red` und jeder Variablen eine `.color-bg`-Klasse zuweist, wobei sich der `color`-Teil für jedes Element in die jeweilige Farbe ändert. Jede Klasse soll auch die `background-color` auf die jeweilige Farbe setzen.

# --hints--

Ihr Code sollte die `@each`-Direktive verwenden.

```js
assert(code.match(/@each /g));
```

Ihre `.blue-bg`-Klasse sollte eine `background-color` in Blau haben.

```js
assert($('.blue-bg').css('background-color') == 'rgb(0, 0, 255)');
```

Ihre `.black-bg`-Klasse sollte eine `background-color` in Schwarz haben.

```js
assert($('.black-bg').css('background-color') == 'rgb(0, 0, 0)');
```

Ihre `.red-bg`-Klasse sollte eine `background-color` in Rot haben.

```js
assert($('.red-bg').css('background-color') == 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```

# --solutions--

```html
<style type='text/scss'>

  @each $color in blue, black, red {
    .#{$color}-bg {background-color: $color;}
  }

  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```

---

```html
<style type='text/scss'>

  $colors: (color1: blue, color2: black, color3: red);

  @each $key, $color in $colors {
    .#{$color}-bg {background-color: $color;}
  }

  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```
