---
id: 587d7dbf367417b2b2512bbb
title: "Einen Stil anwenden, bis eine Bedingung mit @while erfüllt ist"
challengeType: 0
forumTopicId: 301454
dashedName: apply-a-style-until-a-condition-is-met-with-while
---

# --description--

Die `@while`-Direktive ist eine Option mit ähnlicher Funktionalität wie die JavaScript-`while`-Schleife. Sie erstellt CSS-Regeln, bis eine Bedingung erfüllt ist.

Die `@for`-Aufgabe gab ein Beispiel, um ein einfaches Grid-System zu erstellen. Das funktioniert auch mit `@while`.

```scss
$x: 1;
@while $x < 13 {
  .col-#{$x} { width: 100%/12 * $x;}
  $x: $x + 1;
}
```

Definieren Sie zuerst eine Variable `$x` und setzen Sie sie auf 1. Verwenden Sie dann die `@while`-Direktive, um das Grid-System *solange* `$x` kleiner als 13 ist, zu erstellen. Nachdem die CSS-Regel für `width` gesetzt wurde, wird `$x` um 1 erhöht, um eine Endlosschleife zu vermeiden.

# --instructions--

Verwenden Sie `@while`, um eine Reihe von Klassen mit unterschiedlichen `font-sizes` zu erstellen.

Es sollten 5 verschiedene Klassen von `text-1` bis `text-5` sein. Setzen Sie dann `font-size` auf `15px` multipliziert mit der aktuellen Indexnummer. Achten Sie darauf, eine Endlosschleife zu vermeiden!

# --hints--

Ihr Code sollte die `@while`-Direktive verwenden.

```js
assert(code.match(/@while /g));
```

Ihr Code sollte eine Indexvariable verwenden, die bei 1 beginnt.

```js
assert(code.match(/\$.*:\s*?1;/gi));
```

Ihr Code sollte die Zählervariable erhöhen.

```js
assert(code.match(/\$(.*)\s*?:\s*\$\1\s*\+\s*1\s*;/gi));
```

Ihre `.text-1`-Klasse sollte eine `font-size` von `15px` haben.

```js
assert($('.text-1').css('font-size') == '15px');
```

Ihre `.text-2`-Klasse sollte eine `font-size` von `30px` haben.

```js
assert($('.text-2').css('font-size') == '30px');
```

Ihre `.text-3`-Klasse sollte eine `font-size` von `45px` haben.

```js
assert($('.text-3').css('font-size') == '45px');
```

Ihre `.text-4`-Klasse sollte eine `font-size` von `60px` haben.

```js
assert($('.text-4').css('font-size') == '60px');
```

Ihre `.text-5`-Klasse sollte eine `font-size` von `75px` haben.

```js
assert($('.text-5').css('font-size') == '75px');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

# --solutions--

```html
<style type='text/scss'>
  $x: 1;
  @while $x < 6 {
    .text-#{$x}{
      font-size: 15px * $x;
    }
    $x: $x + 1;
  }
</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```
