---
id: 66ed41f912d0bb1dc62da5dd
title: 建置一個 JavaScript 趣味問答機器人
challengeType: 26
dashedName: lab-javascript-trivia-bot
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該將 `"Hello! I'm your coding fun fact guide!"` 記錄到主控台，作為對使用者的問候訊息。
1. 你應該創建三個變數：`botName`、`botLocation` 和 `favoriteLanguage`，分別用來儲存機器人的名稱、來源地點和最喜歡的程式語言。
1. 你應該將 `"My name is (botName) and I live on (botLocation)."` 記錄到主控台。
1. 你應該將 `"My favorite programming language is (favoriteLanguage)."` 記錄到主控台。
1. 你應該使用 `let` 來創建一個 `codingFact` 變數，並指定一個字串，該字串是關於你機器人最喜愛的程式語言的有趣事實，且包含 `favoriteLanguage` 變數的使用。
1. 你應該將 `codingFact` 記錄到主控台。
1. 你應該再次使用 `favoriteLanguage` 變數，將 `codingFact` 變數指定為關於機器人最愛語言的新事實。
1. 你應該再次將 `codingFact` 記錄到主控台。
1. 你應該再次將 `codingFact` 變數指定為關於機器人最愛語言的另一個新事實，並使用 `favoriteLanguage` 變數。
1. 你應該第三次將 `codingFact` 記錄到主控台。
1. 你應該將 `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` 記錄到主控台，作為機器人的告別述語。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該將 `"Hello! I'm your coding fun fact guide!"` 記錄到主控台。

```js
assert.equal(getLogs()[0], "Hello! I'm your coding fun fact guide!")
```

你應該宣告一個 `botName` 變數。請仔細檢查是否有任何拼寫或大小寫錯誤。

```js
assert.exists(botName);
```

你的 `botName` 變數應該是一個字串。

```js
assert.isString(botName);
```

你應該宣告一個 `botLocation` 變數。請仔細檢查是否有任何拼寫或大小寫錯誤。

```js
assert.exists(botLocation);
```

你的 `botLocation` 變數應該是一個字串。

```js
assert.isString(botLocation);
```

你應該宣告一個 `favoriteLanguage` 變數。請再次確認是否有任何拼寫或大小寫錯誤。

```js
assert.exists(favoriteLanguage);
```

你的 `favoriteLanguage` 變數應該是一個字串。

```js
assert.isString(favoriteLanguage);
```

你應該在主控台記錄 `"My name is (botName) and I live on (botLocation)."` 並將變數添加到字串中。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[1], `My name is ${botName} and I live on ${botLocation}.`)
```

你應該在主控台記錄 `"My favorite programming language is (favoriteLanguage)."` 並將變數添加到字串中。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[2], `My favorite programming language is ${favoriteLanguage}.`)
```

你應該使用 `let` 來宣告一個新的變數 `codingFact`。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/)
```

你應該給 `codingFact` 一個包含 `favoriteLanguage` 的值。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/);
assert.match(code, /codingFact\s*=\s*(("|'|`)?.+?\1?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|'|`)?.+?\3?))/);
assert.exists(codingFact);
assert.isNotEmpty(codingFact);
assert.include(String(codingFact), favoriteLanguage);
assert.include(getLogs()[3], favoriteLanguage);
```

你應該將 `codingFact` 記錄到主控台。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[3], favoriteLanguage);
assert.isAtLeast(loggingCodingFacts.length, 1);
```

你應該為 `codingFact` 指定一個也包含 `favoriteLanguage` 的新值，並將它記錄到主控台。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
const [first, second, third] = codeWithoutComments.match(/(let )?\s*codingFact\s*=\s*(("|')?.+?\2?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|')?.+?\2?))/g);
assert.include(getLogs()[4], favoriteLanguage);
assert.notEqual(getLogs()[4], getLogs()[3]);
assert.isAtLeast(loggingCodingFacts.length, 2);
assert.exists(second);
assert.isNotEmpty(codingFact); 
```

你應該為第三次的 `codingFact` 指定一個也包含 `favoriteLanguage` 的值，並將它記錄到主控台。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[5], favoriteLanguage);
assert.notEqual(getLogs()[5], getLogs()[4]);
assert.equal(getLogs()[5], codingFact);
assert.lengthOf(loggingCodingFacts, 3);
assert.isNotEmpty(codingFact); 
```

你應該在主控台記錄 `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` 並添加這些變數的值。

```js
assert.equal(getLogs()[6], `It was fun sharing these facts with you. Goodbye! - ${botName} from ${botLocation}.`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
console.log("Hello! I'm your coding fun fact guide!");

const botName = "JsBot";
const botLocation = "planet Eris";
const favoriteLanguage = "JavaScript";

console.log("My name is " + botName + " and I live on " + botLocation + ".");
console.log("My favorite programming language is " + favoriteLanguage + ".");

let codingFact = "Did you know that " + favoriteLanguage + " was created in just 10 days?";

console.log(codingFact);

codingFact = "Another fun fact: " + favoriteLanguage + "was originally called Mocha!";

console.log(codingFact);

codingFact = "Also, " + favoriteLanguage + " is the most popular programming language in the world.";

console.log(codingFact);

console.log("It was fun sharing these facts with you. Goodbye! - " + botName + " from " + botLocation + ".");

```
