---
id: 587d7fb5367417b2b2512c03
title: 의존성의 최신 마이너 버전을 사용하려면 캐럿 문자를 사용하세요
challengeType: 2
forumTopicId: 301531
dashedName: use-the-caret-character-to-use-the-latest-minor-version-of-a-dependency
---

# --description--

지난 도전 과제에서 배운 물결표가 npm이 의존성의 최신 패치 버전을 설치하도록 허용하는 것과 비슷하게, 캐럿(`^`)은 npm이 향후 업데이트도 설치하도록 허용합니다. 차이점은 캐럿이 마이너 업데이트와 패치 모두를 허용한다는 점입니다.

현재 `@freecodecamp/example`의 버전은 `~1.2.13`여야 하며, 이는 npm이 최신 `1.2.x` 버전으로 설치하도록 허용합니다. 만약 버전 접두사로 캐럿(^)을 사용한다면, npm은 어떤 `1.x.x` 버전으로도 업데이트할 수 있습니다.

```json
"package": "^1.3.8"
```

이렇게 하면 패키지의 어떤 `1.x.x` 버전으로도 업데이트할 수 있습니다.

# --instructions--

의존성에서 `^`의 버전에 캐럿(`@freecodecamp/example`)을 접두사로 사용하여 npm이 새로운 마이너 릴리스로 업데이트할 수 있게 하세요.

**참고:** 버전 번호 자체는 변경하지 마세요.

# --hints--

`"dependencies"`에 `"@freecodecamp/example"`가 포함되어야 합니다.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

`"@freecodecamp/example"` 버전은 `"^1.x.x"`와 일치해야 합니다.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^\^1\./,
        'Wrong version of "@freecodecamp/example". It should be ^1.2.13'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
