---
id: 587d7fb4367417b2b2512bff
title: package.json에 버전 추가하기
challengeType: 2
forumTopicId: 301525
dashedName: add-a-version-to-your-package-json
---

# --description--

`version`는 package.json 파일의 필수 필드 중 하나입니다. 이 필드는 프로젝트의 현재 버전을 설명합니다. 예시는 다음과 같습니다:

```json
"version": "1.2.0",
```

# --instructions--

프로젝트의 package.json 파일에 `version`를 추가하세요.

# --hints--

package.json에는 유효한 "version" 키가 있어야 합니다.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.version, '"version" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
