---
id: 587d7fb5367417b2b2512c02
title: "의존성의 최신 패치 버전을 항상 사용하려면 물결표 문자를 사용하세요"
challengeType: 2
forumTopicId: 301532
dashedName: use-the-tilde-character-to-always-use-the-latest-patch-version-of-a-dependency
---

# --description--

지난 도전 과제에서는 npm에 특정 버전의 패키지만 포함하도록 지시했습니다. 이는 프로젝트의 여러 부분이 서로 호환 가능하도록 유지해야 할 때 의존성을 고정하는 유용한 방법입니다. 하지만 대부분의 경우, 중요한 보안 패치가 포함된 버그 수정이 자주 이루어지므로 이런 수정 사항을 놓치고 싶지 않습니다(운이 좋으면 문제를 일으키지도 않습니다).

npm 의존성이 최신 패치 버전으로 업데이트되도록 허용하려면, 의존성 버전 앞에 물결표(`~`) 문자를 붙일 수 있습니다. 다음은 어떤 `1.3.x` 버전으로든 업데이트를 허용하는 예시입니다.

```json
"package": "~1.3.8"
```

# --instructions--

package.json 파일에서 현재 npm이 `@freecodecamp/example`를 업그레이드하는 규칙은 특정 버전(`1.2.13`)을 사용하는 것입니다. 하지만 이제 최신 `1.2.x` 버전을 허용하려고 합니다.

의존성에 있는 `~` 버전 앞에 물결표(`@freecodecamp/example`) 문자를 붙여 npm이 새로운 _패치_ 릴리스로 업데이트할 수 있도록 하세요.

**참고:** 버전 번호 자체는 변경하지 마세요.

# --hints--

`"dependencies"`에는 `"@freecodecamp/example"`가 포함되어야 합니다.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

`"@freecodecamp/example"` 버전은 `"~1.2.13"`와 일치해야 합니다.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^\~1\.2\.13/,
        'Wrong version of "@freecodecamp/example". It should be ~1.2.13'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
