---
id: 587d7fb4367417b2b2512bfe
title: package.json에 라이선스 추가하기
challengeType: 2
forumTopicId: 301523
dashedName: add-a-license-to-your-package-json
---

# --description--

`license` 필드는 사용자에게 프로젝트에서 무엇을 할 수 있는지 알려주는 곳입니다.

오픈 소스 프로젝트에 흔히 사용되는 라이선스에는 MIT와 BSD가 있습니다. 라이선스 정보는 필수가 아니며, 대부분 국가의 저작권법에 따라 기본적으로 창작물에 대한 소유권이 부여됩니다. 하지만 사용자가 무엇을 할 수 있고 할 수 없는지를 명확히 밝히는 것이 항상 좋은 습관입니다. 다음은 license 필드의 예시입니다:

```json
"license": "MIT",
```

# --instructions--

프로젝트의 package.json 파일에서 `license` 필드를 적절히 채우세요.

# --hints--

package.json에는 유효한 "license" 키가 있어야 합니다.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.license, '"license" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
