---
id: 587d7fb3367417b2b2512bfc
title: package.json에 설명 추가하기
challengeType: 2
forumTopicId: 301522
dashedName: add-a-description-to-your-package-json
---

# --description--

좋은 package.json 파일의 다음 부분은 `description` 필드입니다. 이곳에는 프로젝트에 대한 짧지만 유익한 설명이 들어갑니다.

언젠가 npm에 패키지를 게시할 계획이라면, 사용자가 패키지를 설치할지 결정할 때 이 문자열이 아이디어를 판매하는 역할을 합니다. 하지만 설명의 용도는 그것만이 아닙니다. 프로젝트가 무엇을 하는지 요약하는 훌륭한 방법이기도 합니다. 다른 개발자, 미래의 유지보수 담당자 또는 미래의 자신이 프로젝트를 빠르게 이해하는 데 Node.js 프로젝트에서 매우 중요합니다.

프로젝트 계획과 상관없이 설명은 반드시 포함하는 것이 좋습니다. 예시는 다음과 같습니다.

```json
"description": "A project that does something awesome",
```

# --instructions--

프로젝트의 package.json 파일에 `description`을 추가하세요.

**참고:** 필드 이름에는 반드시 큰따옴표(")를 사용하고, 필드 구분에는 쉼표(,)를 사용하세요.

# --hints--

package.json에 유효한 "description" 키가 있어야 합니다.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.description, '"description" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
