---
id: 587d7fb5367417b2b2512c01
title: 의미 있는 버전 관리를 이해하여 npm 의존성 관리하기
challengeType: 2
forumTopicId: 301529
dashedName: manage-npm-dependencies-by-understanding-semantic-versioning
---

# --description--

`Versions` package.json 파일의 dependencies 섹션에 있는 npm 패키지들은 의미 있는 버전 관리(SemVer)라고 하는 소프트웨어 버전 관리의 업계 표준을 따릅니다. 이는 의존성을 더 쉽게 관리할 수 있도록 하는 것을 목표로 합니다. npm에 게시된 라이브러리, 프레임워크 또는 기타 도구는 프로젝트가 업데이트할 때 어떤 종류의 변경 사항이 예상되는지 명확히 전달하기 위해 SemVer를 사용해야 합니다.

외부 의존성을 사용하는 소프트웨어를 개발할 때(거의 항상 그렇습니다) SemVer를 아는 것이 유용합니다. 언젠가 이 숫자들의 의미를 이해하는 것이 어제는 잘 작동하던 것이 오늘 갑자기 작동하지 않는 이유를 모른 채 프로젝트에 치명적인 변경을 실수로 도입하는 것을 막아줄 것입니다. 공식 웹사이트에 따르면 의미 있는 버전 관리는 다음과 같이 작동합니다:

```json
"package": "MAJOR.MINOR.PATCH"
```

MAJOR 버전은 호환되지 않는 API 변경을 할 때 증가해야 합니다. MINOR 버전은 이전 버전과 호환되는 방식으로 기능을 추가할 때 증가해야 합니다. PATCH 버전은 이전 버전과 호환되는 버그 수정을 할 때 증가해야 합니다. 즉, PATCH는 버그 수정이고 MINOR는 새로운 기능을 추가하지만 둘 다 이전에 작동하던 것을 깨뜨리지 않습니다. 마지막으로 MAJOR는 이전 버전과 호환되지 않는 변경을 추가합니다.

# --instructions--

`package.json` 파일의 dependencies 섹션에서 `@freecodecamp/example`의 버전을 MAJOR 버전 1, MINOR 버전 2, PATCH 버전 13에 맞게 변경하세요.

# --hints--

`"dependencies"`에는 `"@freecodecamp/example"`가 포함되어야 합니다.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

`"@freecodecamp/example"` 버전은 `"1.2.13"`이어야 합니다.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.equal(
        data.dependencies['@freecodecamp/example'],
        '1.2.13',
        'Wrong version of "@freecodecamp/example". It should be 1.2.13'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
