---
id: 587d7fb5367417b2b2512c04
title: 종속성에서 패키지 제거하기
challengeType: 2
forumTopicId: 301530
dashedName: remove-a-package-from-your-dependencies
---

# --description--

이제 package.json의 dependencies 섹션을 사용해 프로젝트의 종속성을 관리하는 몇 가지 방법을 테스트해 보았습니다. 또한 외부 패키지를 파일에 추가하고, 물결표나 캐럿 같은 특수 문자를 사용해 원하는 버전 유형을 npm에 알려주기도 했습니다.

그런데 더 이상 필요하지 않은 외부 패키지를 제거하려면 어떻게 해야 할까요? 이미 짐작했겠지만, 종속성에서 해당 패키지에 해당하는 키-값 쌍을 제거하면 됩니다.

이 방법은 package.json의 다른 필드를 제거할 때도 동일하게 적용됩니다.

# --instructions--

종속성에서 `@freecodecamp/example` 패키지를 제거하세요.

**참고:** 제거 후에 쉼표 개수가 올바른지 꼭 확인하세요.

# --hints--

`"dependencies"`에는 `"@freecodecamp/example"`가 포함되지 않아야 합니다.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.notProperty(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" still includes "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
