---
id: 587d7fb4367417b2b2512c00
title: npm의 외부 패키지로 프로젝트 확장하기
challengeType: 2
forumTopicId: 301527
dashedName: expand-your-project-with-external-packages-from-npm
---

# --description--

패키지 관리자를 사용하는 가장 큰 이유 중 하나는 강력한 의존성 관리입니다. 새 컴퓨터에서 프로젝트를 설정할 때마다 모든 의존성을 수동으로 확인할 필요 없이, npm이 자동으로 모든 것을 설치해 줍니다. 그렇다면 npm이 프로젝트에 정확히 어떤 것이 필요한지 어떻게 알 수 있을까요? package.json 파일의 `dependencies` 섹션을 만나 보세요.

이 섹션에는 프로젝트가 필요로 하는 패키지들이 다음 형식으로 저장됩니다:

```json
"dependencies": {
  "package-name": "version",
  "express": "4.14.0"
}

```

# --instructions--

`1.1.0` 파일의 `@freecodecamp/example` 필드에 `dependencies` 패키지의 `package.json` 버전을 추가하세요.

**참고:** `@freecodecamp/example`는 학습 도구로 사용되는 가짜 패키지입니다.

# --hints--

`"dependencies"`에는 `"@freecodecamp/example"`가 포함되어야 합니다.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

`"@freecodecamp/example"` 버전은 `"1.1.0"`이어야 합니다.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^[\^\~]?1\.1\.0/,
        'Wrong version of "@freecodecamp/example" installed. It should be 1.1.0'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
