---
id: 587d7fb4367417b2b2512bfd
title: package.json에 키워드 추가하기
challengeType: 2
forumTopicId: 301526
dashedName: add-keywords-to-your-package-json
---

# --description--

`keywords` 필드는 관련 키워드를 사용해 프로젝트를 설명할 수 있는 곳입니다. 예시는 다음과 같습니다:

```json
"keywords": [ "descriptive", "related", "words" ],
```

보시다시피, 이 필드는 큰따옴표로 묶인 문자열 배열로 구성되어 있습니다.

# --instructions--

프로젝트의 package.json 파일에서 `keywords` 필드에 적절한 문자열 배열을 추가하세요.

키워드 중 하나는 "freecodecamp"여야 합니다.

# --hints--

package.json에 유효한 "keywords" 키가 있어야 합니다

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.keywords, '"keywords" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

"keywords" 필드는 배열이어야 합니다

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.isArray(data.keywords, '"keywords" is not an array');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

"keywords"에 "freecodecamp"가 포함되어야 합니다

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.include(
        data.keywords,
        'freecodecamp',
        '"keywords" does not include "freecodecamp"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
