---
id: 587d7fb3367417b2b2512bfb
title: 'Node.js 프로젝트나 npm 패키지의 핵심인 package.json 사용법'
challengeType: 2
forumTopicId: 301528
dashedName: how-to-use-package-json-the-core-of-any-node-js-project-or-npm-package
---

# --description--

이 도전 과제들은 다음 방법 중 하나를 사용해 코드를 작성하는 것을 포함합니다:

- <a href="https://github.com/freeCodeCamp/boilerplate-npm/" target="_blank" rel="noopener noreferrer nofollow">이 GitHub 저장소</a>를 클론하여 로컬에서 도전 과제를 완료하세요.
- 원하는 사이트 빌더를 사용해 프로젝트를 완료하세요. 반드시 GitHub 저장소의 모든 파일을 포함해야 합니다.

`package.json` 파일은 모든 Node.js 프로젝트나 npm 패키지의 중심입니다. 이 파일은 프로젝트에 관한 정보를 저장합니다. 정보는 키-값 쌍으로 저장되는 단일 JSON 객체로 구성됩니다. 필수 필드는 `name`과 `version` 두 개뿐이지만, 추가 정보를 제공하는 것이 좋은 습관입니다.

터미널에서 `package.json` 명령어를 사용해 `npm init` 파일을 생성할 수 있습니다. 이 명령어는 안내 설정을 실행합니다. `npm init`에 `-y` 플래그를 붙여 실행하면 질문 없이 파일을 생성할 수 있습니다, `npm init -y`.

프로젝트의 파일 트리를 보면 최상위에 `package.json` 파일이 있습니다. 이 파일을 다음 도전 과제들에서 개선할 것입니다.

이 파일에서 가장 흔히 포함되는 정보 중 하나는 `author` 필드입니다. 이 필드는 프로젝트를 만든 사람을 지정하며, 문자열이나 연락처 등 다른 세부 정보를 담은 객체일 수 있습니다. 큰 프로젝트에는 객체를 권장하지만, 이 프로젝트에는 다음 예시처럼 간단한 문자열이면 충분합니다.

```json
"author": "Jane Doe",
```

# --instructions--

`author` 파일에 프로젝트의 `package.json`로 본인 이름을 추가하세요.

**참고:** JSON을 작성하는 것이므로 모든 필드명은 큰따옴표(")로 감싸고 쉼표(,)로 구분해야 합니다.

# --hints--

`package.json`에는 유효한 "author" 키가 있어야 합니다

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.author, '"author" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
