---
id: 685dbd19200ad656d48b77ff
title: بناء صفحة عربة تسوق موسيقية
challengeType: 25
dashedName: lab-music-shopping-cart-page
demoType: onClick
---

# --description--

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك عنصر يحمل `id` باسم `shopping-cart-container`.
2. يجب أن يستخدم العنصر `#shopping-cart-container` صنف الأدوات المناسب لتعيين العنصر إلى تخطيط Flexbox.
3. يجب أن يضبط العنصر `#shopping-cart-container` اتجاه العناصر المرنة إلى عمود للأجهزة الصغيرة وإلى صف للأجهزة الأكبر. تذكر أن Tailwind CSS تعتمد نهج الأجهزة المحمولة أولاً، وستستخدم بادئة نقطة التوقف `lg:` لاستهداف الأجهزة الأكبر.
4. داخل العنصر `#shopping-cart-container`، يجب أن يكون لديك عنصر يحمل `id` باسم `products-container`.
5. يجب أن يحتوي العنصر `#products-container` على عنصرين فرعيين على الأقل، كل منهما يحمل صنفًا باسم `card`.
6. يجب أن يحتوي كل عنصر `.card` على العناصر التالية متداخلة بداخله:
   - عنصر `h2` يحتوي على نص يمثل اسم المنتج، ويستخدم صنف أدوات يضبط حجم الخط باستخدام أصناف الحجم المعرفة مسبقًا فقط مثل `text-sm`، `text-md`، `text-lg`، `text-xl`، `text-2xl`، وغيرها.
   - عنصر يحمل صنفًا باسم `quantity` ونصًا يمثل عدد عناصر السلة لذلك المنتج.
   - عنصر يحمل صنفًا باسم `price` ونصًا يمثل السعر.
   - زر `button` يحمل صنفًا باسم `remove-button` ونصًا `Remove`. يجب أن يحتوي زر `button` على أصناف أدوات لخلفية حمراء معرفة مسبقًا تختارها، وأصناف أدوات مختلفة لخلفية حمراء عند حالة التحويم. من أمثلة الخلفيات الحمراء المعرفة مسبقًا `bg-red-500`، `bg-red-600`، وغيرها. ومن أمثلة الخلفيات الحمراء عند التحويم `hover:bg-red-500`، `hover:bg-red-600`، وغيرها.
7. داخل العنصر `#shopping-cart-container`، يجب أن يكون لديك عنصر يحمل `id` باسم `order-summary-container`.
8. يجب أن يحتوي العنصر `#order-summary-container` على الأنماط التالية:
   - صنف أدوات تختاره لتعيين زوايا مستديرة معرفة مسبقًا. من الأمثلة `rounded`، `rounded-lg`، `rounded-full`، وغيرها.
   - صنف أدوات تختاره لتعيين عرض الحدود على جميع الجوانب.
9. يجب أن يحتوي العنصر `#order-summary-container` على العناصر التالية متداخلة بداخله:
   - عنصر `h2` يحتوي على النص `Order Summary` وصنف أدوات تختاره يضبط حجم الخط.
   - عنصر يحتوي على النص `Total:` ويحمل `id` باسم `total`. يجب أن يستخدم هذا العنصر أصناف أدوات لضبط وزن الخط إلى وزن خط معرف مسبقًا في Tailwind CSS وحجم الخط الذي تختاره.
   - عنصر يحمل `id` باسم `total-amount` لعرض المجموع الكلي لجميع العناصر في السلة.
   - رابط يحتوي على النص `Checkout` وصنف الأدوات الصحيح لمركزة النص. يجب أن تكون قيمة `href` مضبوطة على `#`. يجب أن يحتوي الرابط أيضًا على صنف أدوات لتعيين لون الخلفية إلى لون أزرق معرف مسبقًا في Tailwind CSS تختاره، ولون أزرق مختلف لخلفية التحويم.

# --hints--

يجب أن يكون لديك عنصر يحمل `id` مضبوطًا على `shopping-cart-container`.

```js
assert.exists(document.getElementById("shopping-cart-container"));
```

يجب أن يستخدم العنصر `#shopping-cart-container` الصنف المساعد `flex` لضبط العنصر على تخطيط Flexbox.

```js
const shoppingCartContainer = document.getElementById("shopping-cart-container");
assert.isTrue(shoppingCartContainer.classList.contains("flex"));
```

يجب أن يستخدم العنصر `#shopping-cart-container` الصنف المساعد الصحيح لضبط اتجاه العناصر المرنة إلى عمود.

```js
const shoppingCartContainer = document.getElementById("shopping-cart-container");
assert.isTrue(shoppingCartContainer.classList.contains("flex-col"));
```

يجب أن يستخدم العنصر `#shopping-cart-container` بادئة نقطة التوقف `lg:` متبوعة بالصنف المساعد الصحيح لضبط اتجاه العناصر المرنة إلى صف للأجهزة الأكبر.

```js
const shoppingCartContainer = document.getElementById("shopping-cart-container");
assert.isTrue(shoppingCartContainer.classList.contains("lg:flex-row"));
```

يجب أن يكون لديك عنصر يحمل `id` مضبوطًا على `products-container` داخل العنصر `#shopping-cart-container`.

```js
const shoppingCartContainer = document.getElementById("shopping-cart-container");
assert.exists(shoppingCartContainer);
const productsContainer = document.getElementById("products-container");
assert.exists(productsContainer);
assert.isTrue(shoppingCartContainer.contains(productsContainer));
```

يجب أن يحتوي العنصر `#products-container` على عنصرين فرعيين على الأقل لكل منهما صنف يسمى `card`.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");
assert.isAtLeast(cards.length, 2);
```

يجب أن يحتوي كل عنصر `.card` على عنصر `h2` يحتوي على نص يمثل اسم المنتج.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");

cards?.forEach((card) => {
  const productName = card.querySelector("h2");
  assert.exists(productName);
  assert.isNotEmpty(productName.textContent);
});
```

يجب أن يستخدم كل عنصر `h2` داخل عنصر `.card` صنفًا مساعدًا يضبط حجم الخط باستخدام أصناف الحجم المعرفة مسبقًا فقط مثل `text-sm`، `text-base`، `text-lg`، `text-xl`، `text-2xl`، إلخ.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

cards?.forEach((card) => {
  const productName = card.querySelector("h2");
  assert.exists(productName);
  const hasPredefinedSize = predefinedSizes.some(size => productName.classList.contains(size));
  assert.isTrue(hasPredefinedSize);
});
```

يجب أن يحتوي كل عنصر `.card` على عنصر يحمل صنفًا يسمى `quantity`.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");

cards?.forEach((card) => {
  const quantity = card.querySelector(".quantity");
  assert.exists(quantity);
});
```

يجب أن يحتوي كل عنصر `.quantity` على نص يمثل عدد عناصر العربة لذلك المنتج.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");

cards?.forEach((card) => {
  const quantity = card.querySelector(".quantity");
  assert.exists(quantity);
  assert.isNotEmpty(quantity.textContent.trim());
});
```

يجب أن يحتوي كل عنصر `.card` على عنصر يحمل صنفًا يسمى `price`.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");  

cards?.forEach((card) => {
  const price = card.querySelector(".price");
  assert.exists(price);
});
```

يجب أن يحتوي كل عنصر `.price` على نص يمثل السعر.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");

cards?.forEach((card) => {
  const price = card.querySelector(".price");
  assert.exists(price);
  assert.isNotEmpty(price.textContent.trim());
});
```

يجب أن يحتوي كل عنصر من نوع `.card` على عنصر `button` يحمل الصنف `remove-button` ونصه `Remove`.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");

cards?.forEach((card) => {
  const removeButton = card.querySelector("button.remove-button");
  assert.exists(removeButton);
  assert.equal(removeButton.textContent.trim(), "Remove");
});
```

يجب أن يحتوي كل عنصر `button` على فئة مساعدة لخلفية حمراء تختارها.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");
const predefinedRedColors = [
  "bg-red-50",
  "bg-red-100",
  "bg-red-200",
  "bg-red-300",
  "bg-red-400",
  "bg-red-500",
  "bg-red-600",
  "bg-red-700",
  "bg-red-800",
  "bg-red-900",
  "bg-red-950"
];

cards?.forEach((card) => {
  const removeButton = card.querySelector("button.remove-button");
  assert.exists(removeButton);
  const hasRedBackground = predefinedRedColors.some(color => removeButton.classList.contains(color));
  assert.isTrue(hasRedBackground);
});
```

يجب أن يحتوي كل عنصر `button` على لون خلفية أحمر مختلف لحالة التحويم.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");

const predefinedRedHoverColors = [
  "hover:bg-red-50",
  "hover:bg-red-100",
  "hover:bg-red-200",
  "hover:bg-red-300",
  "hover:bg-red-400",
  "hover:bg-red-500",
  "hover:bg-red-600",
  "hover:bg-red-700",
  "hover:bg-red-800",
  "hover:bg-red-900",
  "hover:bg-red-950"
];

cards?.forEach((card) => {
  const removeButton = card.querySelector("button.remove-button");
  assert.exists(removeButton);

  const classList = removeButton.classList;

  const hoverBgClass = [...classList].find(cls => cls.startsWith("hover:bg-red-"));
  const bgClass = [...classList].find(cls => cls.startsWith("bg-red-") && !cls.startsWith("hover:"));

  assert.notEqual(bgClass, hoverBgClass.replace("hover:", ""));
});
```

يجب أن يكون لديك عنصر يحمل `id` يسمى `order-summary-container` داخل العنصر `#shopping-cart-container`.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
assert.exists(orderSummaryContainer);
```

يجب أن يحتوي العنصر `#order-summary-container` على صنف مساعد تختاره لضبط الأركان المدورة.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const predefinedRoundedCorners = [
  "rounded-none",
  "rounded-sm",
  "rounded",
  "rounded-md",
  "rounded-lg",
  "rounded-xl",
  "rounded-2xl",
  "rounded-3xl",
  "rounded-full"
];

const hasRoundedCorners = predefinedRoundedCorners.some(corner => orderSummaryContainer.classList.contains(corner));
assert.isTrue(hasRoundedCorners);
```

يجب أن يحتوي العنصر `#order-summary-container` على صنف مساعد تختاره لضبط عرض الحدود على جميع الجوانب.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");

// border classes can vary in width, so this check is for the presence of any border class
const hasBorder = [...orderSummaryContainer.classList].some(className => className.startsWith("border-") || className === "border");
assert.isTrue(hasBorder);
```

يجب أن يحتوي العنصر `#order-summary-container` على عنصر `h2` يحتوي على النص `Order Summary`.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const orderSummaryTitle = orderSummaryContainer.querySelector("h2");
assert.exists(orderSummaryTitle);
assert.equal(orderSummaryTitle.textContent.trim(), "Order Summary");
```

يجب أن يستخدم عنصر `h2` داخل العنصر `#order-summary-container` صنفًا مساعدًا تختاره يضبط حجم الخط.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

const orderSummaryTitle = orderSummaryContainer.querySelector("h2");
assert.exists(orderSummaryTitle);
const hasPredefinedSize = predefinedSizes.some(size => orderSummaryTitle.classList.contains(size));
assert.isTrue(hasPredefinedSize);
```

يجب أن يحتوي العنصر `#order-summary-container` على عنصر يحتوي على النص `Total:` ويحمل `id` مضبوطًا على `total`.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const totalElement = orderSummaryContainer.querySelector("#total");
assert.exists(totalElement);
assert.equal(totalElement.textContent.trim(), "Total:");
```

يجب أن يستخدم العنصر `#total` صنفًا مساعدًا لضبط وزن الخط.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const totalElement = orderSummaryContainer.querySelector("#total");
const predefinedFontWeights = [
  "font-thin",
  "font-extralight",
  "font-light",
  "font-normal",
  "font-medium",
  "font-semibold",
  "font-bold",
  "font-extrabold",
  "font-black"
];

const hasFontWeight = predefinedFontWeights.some(weight => totalElement.classList.contains(weight));
assert.isTrue(hasFontWeight);
```

يجب أن يستخدم العنصر `#total` صنفًا مساعدًا لضبط حجم الخط.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const totalElement = orderSummaryContainer.querySelector("#total");
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

const hasPredefinedSize = predefinedSizes.some(size => totalElement.classList.contains(size));
assert.isTrue(hasPredefinedSize);
```

يجب أن يحتوي العنصر `#order-summary-container` على عنصر يحمل `id` مضبوطًا على `total-amount` لعرض المجموع لجميع العناصر في العربة.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const totalElement = orderSummaryContainer.querySelector("#total-amount");
assert.exists(totalElement);
assert.isNotEmpty(totalElement.textContent.trim());
```

يجب أن يحتوي العنصر `#order-summary-container` على رابط يحتوي على النص `Checkout` وقيمة `href` مضبوطة على `#`.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const checkoutLink = orderSummaryContainer.querySelector("a");
assert.exists(checkoutLink);
assert.equal(checkoutLink.textContent.trim(), "Checkout");
assert.equal(checkoutLink.getAttribute("href"), "#");
```

يجب أن يحتوي رابط `Checkout` على صنف مساعد لمركزة النص.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const checkoutLink = orderSummaryContainer.querySelector("a");
const hasTextCentering = checkoutLink.classList.contains("text-center");
assert.isTrue(hasTextCentering);
```

يجب أن يحتوي رابط `Checkout` على صنف مساعد لضبط لون الخلفية إلى لون أزرق معرف مسبقًا في Tailwind CSS تختاره.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const checkoutLink = orderSummaryContainer.querySelector("a");
const predefinedBlueColors = [
  "bg-blue-50",
  "bg-blue-100",
  "bg-blue-200",
  "bg-blue-300",
  "bg-blue-400",
  "bg-blue-500",
  "bg-blue-600",
  "bg-blue-700",
  "bg-blue-800",
  "bg-blue-900",
  "bg-blue-950"
];

const hasBlueBackground = predefinedBlueColors.some(color => checkoutLink.classList.contains(color));
assert.isTrue(hasBlueBackground);
```

يجب أن يحتوي رابط `Checkout` على لون خلفية أزرق مختلف للحالة عند المرور.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const checkoutLink = orderSummaryContainer.querySelector("a");
const predefinedBlueHoverColors = [
  "hover:bg-blue-50",
  "hover:bg-blue-100",
  "hover:bg-blue-200",
  "hover:bg-blue-300",
  "hover:bg-blue-400",
  "hover:bg-blue-500",
  "hover:bg-blue-600",
  "hover:bg-blue-700",
  "hover:bg-blue-800",
  "hover:bg-blue-900",
  "hover:bg-blue-950"
];

const classList = checkoutLink.classList;
const hoverBgClass = [...classList].find(cls => cls.startsWith("hover:bg-blue-"));
const bgClass = [...classList].find(cls => cls.startsWith("bg-blue-") && !cls.startsWith("hover:"));

assert.notEqual(bgClass, hoverBgClass.replace("hover:", ""));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music Shopping Cart Page</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>

  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music Shopping Cart Page</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-gray-900 text-white">
    <h1 class="text-center text-3xl my-8">Shopping Cart Page</h1>
    <main
      id="shopping-cart-container"
      class="flex flex-col lg:flex-row items-center gap-8"
    >
      <div id="products-container" class="w-full md:w-2/3 px-4">
        <div
          class="card flex flex-col lg:flex-row items-center gap-6 rounded-lg border border-gray-200 p-4 my-8 bg-gray-800"
        >
          <h2 class="text-lg lg:w-1/2">Ludwig Supraphonic Snare Drum Chrome</h2>
          <div class="btn-container flex items-center lg:w-1/4">
            <button
              class="decrement-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Decrease quantity"
            >
              <span aria-hidden="true">-</span>
            </button>
            <span class="quantity mx-2">2</span>
            <button
              class="increment-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Increase quantity"
            >
              <span aria-hidden="true">+</span>
            </button>
          </div>
          <p class="price lg:w-1/6 text-right">$1398.00</p>
          <button
            class="remove-button bg-red-700 hover:bg-red-800 py-1 px-2 rounded-md"
          >
            Remove
          </button>
        </div>

        <div
          class="card flex flex-col lg:flex-row items-center gap-6 rounded-lg border border-gray-200 p-4 my-8 bg-gray-800"
        >
          <h2 class="text-lg lg:w-1/2">Shure SM58S Mic With Switch Standard</h2>
          <div class="btn-container flex items-center lg:w-1/4">
            <button
              class="decrement-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Decrease quantity"
            >
              <span aria-hidden="true">-</span>
            </button>
            <span class="quantity mx-2">1</span>
            <button
              class="increment-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Increase quantity"
            >
              <span aria-hidden="true">+</span>
            </button>
          </div>
          <p class="price lg:w-1/6 text-right">$119.00</p>
          <button
            class="remove-button bg-red-700 hover:bg-red-800 py-1 px-2 rounded-md"
          >
            Remove
          </button>
        </div>

        <div
          class="card flex flex-col lg:flex-row items-center gap-6 rounded-lg border border-gray-200 p-4 my-8 bg-gray-800"
        >
          <h2 class="text-lg lg:w-1/2">
            Fender American Professional II Stratocaster
          </h2>
          <div class="btn-container flex items-center lg:w-1/4">
            <button
              class="decrement-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Decrease quantity"
            >
              <span aria-hidden="true">-</span>
            </button>
            <span class="quantity mx-2">1</span>
            <button
              class="increment-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Increase quantity"
            >
              <span aria-hidden="true">+</span>
            </button>
          </div>
          <p class="price lg:w-1/6 text-right">$1700.00</p>
          <button
            class="remove-button bg-red-700 hover:bg-red-800 py-1 px-2 rounded-md"
          >
            Remove
          </button>
        </div>

        <div
          class="card flex flex-col lg:flex-row items-center gap-6 rounded-lg border border-gray-200 p-4 my-8 bg-gray-800"
        >
          <h2 class="text-lg lg:w-1/2">
            Moog Subsequent 37 Analog Synthesizer
          </h2>
          <div class="btn-container flex items-center lg:w-1/4">
            <button
              class="decrement-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Decrease quantity"
            >
              <span aria-hidden="true">-</span>
            </button>
            <span class="quantity mx-2">1</span>
            <button
              class="increment-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Increase quantity"
            >
              <span aria-hidden="true">+</span>
            </button>
          </div>
          <p class="price lg:w-1/6 text-right">$1799.00</p>
          <button
            class="remove-button bg-red-700 hover:bg-red-800 py-1 px-2 rounded-md"
          >
            Remove
          </button>
        </div>
      </div>

      <div
        id="order-summary-container"
        class="rounded-lg border border-gray-200 p-8 bg-gray-800 w-full md:w-1/3 mt-8 lg:mt-0"
      >
        <h2 class="text-2xl mb-4">Order Summary</h2>
        <ul>
          <li>
            <span id="subtotal" class="font-semibold text-lg">Subtotal:</span>
            <span>$5,016.00</span>
          </li>
          <li class="my-6">
            <span id="shipping" class="font-semibold text-lg">Shipping:</span>
            <span>Calculated in checkout</span>
          </li>
          <li>
            <span id="total" class="font-semibold text-lg">Total:</span>
            <span id="total-amount">$5,016.00</span>
          </li>
        </ul>
        <a
          href="#"
          class="text-center rounded-lg mx-auto mt-4 block px-5 py-2.5 font-medium bg-blue-600 hover:bg-blue-700"
          >Checkout</a
        >
      </div>
    </main>
  </body>
</html>
```
