---
id: a9bd25c716030ec90084d8a1
title: Chunky Monkeyアルゴリズムを実装する
challengeType: 26
dashedName: implement-the-chunky-monkey-algorithm
---

# --description--

**目的：** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー：**

1. `chunkArrayInGroups` という名前の関数を作成してください。この関数は最初の引数に配列を、2番目の引数に数値を受け取ります。関数は配列を2番目の引数の長さの小さな配列に分割し、それらを二次元配列として返す必要があります。

# --hints--

`chunkArrayInGroups(["a", "b", "c", "d"], 2)` は `[["a", "b"], ["c", "d"]]` を返すべきです。

```js
assert.deepEqual(chunkArrayInGroups(['a', 'b', 'c', 'd'], 2), [
  ['a', 'b'],
  ['c', 'd']
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5], 3)` は `[[0, 1, 2], [3, 4, 5]]` を返すべきです。

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5], 3), [
  [0, 1, 2],
  [3, 4, 5]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5], 2)` は `[[0, 1], [2, 3], [4, 5]]` を返すべきです。

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5], 2), [
  [0, 1],
  [2, 3],
  [4, 5]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5], 4)` は `[[0, 1, 2, 3], [4, 5]]` を返すべきです。

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5], 4), [
  [0, 1, 2, 3],
  [4, 5]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6], 3)` は `[[0, 1, 2], [3, 4, 5], [6]]` を返すべきです。

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6], 3), [
  [0, 1, 2],
  [3, 4, 5],
  [6]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 4)` は `[[0, 1, 2, 3], [4, 5, 6, 7], [8]]` を返すべきです。

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 4), [
  [0, 1, 2, 3],
  [4, 5, 6, 7],
  [8]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 2)` は `[[0, 1], [2, 3], [4, 5], [6, 7], [8]]` を返すべきです。

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 2), [
  [0, 1],
  [2, 3],
  [4, 5],
  [6, 7],
  [8]
]);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function chunkArrayInGroups(arr, size) {
  let out = [];

  for (let i = 0; i < arr.length; i += size) {
    out.push(arr.slice(i, i + size));
  }

  return out;
}

chunkArrayInGroups(['a', 'b', 'c', 'd'], 2);
```
