---
id: 686b9720ee1d032bd77a480a
title: 偵錯 ISBN 驗證器
challengeType: 27
dashedName: debug-an-isbn-validator
---

# --description--

ISBN（國際標準書號）是分配給商業書籍的唯一識別符號。它的長度為 10 或 13 個字元。最後一個字元是檢查碼，該檢查碼是根據前面的數字計算出來的，這些數字稱為主要字元。

Camperbot 嘗試建置他們自己的 ISBN 驗證器。然而，他們在過程中犯了幾個錯誤。

在這個實驗中，你將修正現有的程式碼並使其正常運作。

預期行為：
當使用者執行程式時，會顯示提示 `Enter ISBN and length: `。使用者可以以 `ISBN,length` 格式輸入他們想要驗證的 ISBN 碼。輸入 ISBN 時不包含連字號，接著輸入逗號和其長度（`10` 或 `13`）。

範例輸入：`1530051126,10` 用於 ISBN-10 碼。`9781530051120,13` 用於 ISBN-13 碼。

如何尋找檢查碼：

- 主要字元是 ISBN-10 碼的前 9 個字元，或 ISBN-13 碼的前 12 個字元。它們必須全部是數字。
- 最後一個字元是給定的檢查碼。預期的檢查碼是根據主要字元計算出來的，只有當兩者相符時，ISBN 碼才有效。
- 例如，在 ISBN-10 碼 `1530051126` 中，主要字元是 `153005112`，給定的檢查碼是 `6`。在 ISBN-13 碼 `9781530051120` 中，主要字元是 `978153005112`，給定的檢查碼是 `0`。
- 你不需要知道詳細的計算邏輯。本實驗中的函式（程式）`calculate_check_digit_10` 和 `calculate_check_digit_13` 會負責計算並以字串形式傳回預期的檢查碼。
- ISBN-10 碼的檢查碼可以是從 `0` 到 `9` 的數字或大寫字母 `X`。
- ISBN-13 碼的檢查碼可以是從 `0` 到 `9` 的數字。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該修正目前程式碼中的 `IndentationError`。
1. 即使使用者沒有輸入以逗號分隔的值，程式也應該能處理 `IndexError` 而不會崩潰。
1. 當使用者沒有輸入以逗號分隔的值時，他們應該在主控台看到訊息 `Enter comma-separated values.`，並且程式應該終止。
1. 即使使用者輸入非數字的長度值，程式也應該能處理 `ValueError` 而不會崩潰。
1. 當使用者輸入非數字的長度值時，他們應該在主控台看到訊息 `Length must be a number.`，並且程式應該終止。
1. 你應該修正 `validate_isbn` 函式（程式）中的越界錯誤。
1. 你應該修正目前程式碼中使用者輸入有效 ISBN 碼時發生的 `TypeError`。
1. 你應該修正目前程式碼中使用者輸入有效 ISBN 碼時發生的 `IndexError`。
1. 即使使用者輸入的 ISBN 碼主要字元包含非數字字元，程式也應該能處理發生的 `ValueError` 而不會崩潰。
1. 當使用者輸入的 ISBN 碼主要字元包含非數字字元時，他們應該在主控台看到訊息 `Invalid character was found.`。
1. 當使用者輸入 `1530051126,10` 時，他們應該看到訊息 `Valid ISBN Code.`。
1. 當使用者輸入 `9781530051120,13` 時，他們應該看到訊息 `Valid ISBN Code.`。

**重要**：你需要在全域空間中將 `main()` 呼叫註解掉，測試才能正確執行。

當你完成專案時，使用者應該根據他們輸入的值看到以下訊息。

| ISBN 碼 | 長度 | 訊息 | 範例輸入 |
|-----------|--------|---------| ------------- |
| 有效 | 有效 | `Valid ISBN Code.` | `1530051126,10` |
| 無效數字 | 有效 | `Invalid ISBN Code.` | `1530051125,10` |
| 與指定長度不符或留空 | 有效 | `ISBN-10 code should be 10 digits long.` 或 `ISBN-13 code should be 13 digits long.`，視輸入的長度而定。 | `9781530051120,10` 或 `1530051126,13` |
| 主要字元包含非數字字元 | 有效 | `Invalid character was found.` | `15-0051126,10` |
| 任意 | 無效數字 | `Length should be 10 or 13.` | `1530051126,9` |
| 任意 | 包含非數字字元或留空 | `Length must be a number.` | `1530051125,A` |
| 非逗號分隔 | 非逗號分隔 | `Enter comma-separated values.` | `1530051125` |

如果你想手動測試你的程式碼，可以使用以下值。

有效的 ISBN-10 碼範例輸入：
`1530051126,10`
`9971502100,10`
`080442957X,10`

有效的 ISBN-13 碼範例輸入：
`9781530051120,13`
`9781947172104,13`

# --hints--

你應該將對 `main` 函式（程式）的呼叫註解掉，以允許其餘的測試正常運作。

```js
({
  test: () => runPython('assert not _Node(_code).has_call("main()")')
})
```

你應該有一個 `validate_isbn` 函式（程式）。

```js
({
  test: () => runPython('assert _Node(_code).has_function("validate_isbn")')
})
```

你應該有一個 `calculate_check_digit_10` 函式（程式）。

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_10")')
})
```

你應該有一個 `calculate_check_digit_13` 函式（程式）。

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_13")')
})
```

你應該有一個 `main` 函式（程式）。

```js
({
  test: () => runPython('assert _Node(_code).has_function("main")')
})
```

當使用者輸入非逗號分隔的值時，你應該在主控台看到訊息 `Enter comma-separated values.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

當使用者輸入非數字的長度值時，你應該會在主控台看到訊息 `Length must be a number.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,a"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

當使用者輸入的 ISBN 碼主要字元包含非數字字元時，你應該在主控台看到訊息 `Invalid character was found.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "978IS3OOSII2O,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

當使用者輸入 `1530051126,10` 時，你應該在主控台看到訊息 `Valid ISBN Code.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

當使用者輸入 `9781530051120,13` 時，你應該會看到訊息 `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

當使用者輸入 `1530051125,10` 時，你應該看到訊息 `Invalid ISBN Code.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid ISBN Code." in out

`)
}})
```

當使用者輸入 `9781530051120,10` 時，你應該會看到訊息 `ISBN-10 code should be 10 digits long.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-10 code should be 10 digits long." in out

`)
}})
```

當使用者輸入 `1530051126,13` 時，你應該會看到訊息 `ISBN-13 code should be 13 digits long.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-13 code should be 13 digits long." in out

`)
}})
```

當使用者輸入 `15-0051126,10` 時，你應該會看到訊息 `Invalid character was found.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "15-0051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

當使用者輸入 `1530051126,9` 時，你應該會看到訊息 `Length should be 10 or 13.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,9"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length should be 10 or 13." in out

`)
}})
```

當使用者輸入 `1530051125,A` 時，你應該會看到訊息 `Length must be a number.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,A"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

當使用者輸入 `1530051125` 時，你應該會看到訊息 `Enter comma-separated values.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

當使用者輸入 `9971502100,10` 時，你應該會看到訊息 `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9971502100,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

當使用者輸入 `080442957X,10` 時，你應該會看到訊息 `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "080442957X,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

當使用者輸入 `9781947172104,13` 時，你應該會看到訊息 `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781947172104,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

# --seed--

## --seed-contents--

```py
def validate_isbn(isbn, length):
    if len(isbn, length) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return
    main_digits = isbn[0:length]
    given_check_digit = isbn[length]
    main_digits_list = [int(digit) for digit in main_digits]
    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)
    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')
def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11
    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10
    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]
    length = int(values[1])
    if length == 10 or length == 13:
    validate_isbn(isbn, length)
    else:
    print('Length should be 10 or 13.')

main()
```

# --solutions--

```py
def validate_isbn(isbn, length):
    if len(isbn) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return

    check_digit_index = length - 1

    main_digits = isbn[0:check_digit_index]
    given_check_digit = isbn[check_digit_index]

    try:
        main_digits_list = [int(digit) for digit in main_digits]
    except ValueError as e:
        print('Invalid character was found.')
        return

    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)

    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')


def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11

    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10

    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]

    try:
        length = int(values[1])
    except ValueError as e:
        print('Length must be a number.')
        return
    except IndexError as e:
        print('Enter comma-separated values.')
        return

    if length == 10 or length == 13:
        validate_isbn(isbn, length)
    else:
        print('Length should be 10 or 13.')


# main()
```
