---
id: 6852dd8486c7ee0171322950
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Utahitaji kigezo cha kuhifadhi msimbo wa siri wakati unaufasiri, kwa hivyo ondoa `pass` na unda kigezo kinachoitwa `secret_code` ndani ya kitendakazi na ukaweke thamani ya mfuatano wa herufi tupu.

# --hints--

Unapaswa kuwa na kigezo cha `secret_code` ndani ya kitendakazi cha `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_code')
    `)
})
```

Kigezo cha `secret_code` kinapaswa kuwa na thamani ya mfuatano wa herufi tupu.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_code').is_equivalent('secret_code = ""')
    `)
})
```

Haupasi kuwa na `pass` tena katika msimbo wako.

```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').has_pass()
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
