---
id: 6852db755c9c1274d75aae9a
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika warsha hii utaunda kivutio cha nambari za siri, nambari za siri zimetoweka katika kila mstari wa shairi.

Anza kwa kuunda kitendakazi kinachoitwa `pin_extractor` chenye kigezo `poem`. Kitendakazi kisicho na chochote ndani kitatoa kosa, kwa hivyo ongeza `pass` ndani ya kitendakazi kwa muda ili kiwe halali.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('pin_extractor')
    `)
})
```

Kitendakazi `pin_extractor` kinapaswa kuwa na kigezo kinachoitwa `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poem')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
