---
id: 68530231c05b93324e071104
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Wakati unafanya kazi kwenye kitendakazi, itakuwa na manufaa kuwa na wito wa kitendakazi kwa ajili ya kurekebisha makosa, kwa hivyo tengeneza kigezo cha kimataifa `poem` nje ya kitendakazi, chenye shairi hili:

```md
Stars and the moon
shine in the sky
white and bright
until the end of the night
```

Unaweza kutumia mfuatano wa herufi wa mistari mingi kwa ajili hiyo.

Kisha, fanya wito wa kitendakazi kwa kutumia kigezo `poem` kama hoja.

# --hints--

Unapaswa kuwa na kigezo cha kimataifa kinachoitwa `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_variable('poem')
    `)
})
```

Kigezo `poem` kinapaswa kuwa na thamani sawa na `Stars and the moon\nshine in the sky\nwhite and bright\nuntil the end of the night`.

```js
({
    test: () => runPython(`
    pm = """Stars and the moon
    shine in the sky
    white and bright
    until the end of the night"""
    assert _Node(_code).find_variable('poem').is_equivalent(f'poem = """{pm}"""')
    `)
})
```

Unapaswa kuita kitendakazi `pin_extractor` kwa kutumia `poem` kama hoja.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('pin_extractor(poem)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''

--fcc-editable-region--

--fcc-editable-region--
```
