---
id: 68531fbb934d735415053274
title: Hatua ya 4
challengeType: 20
dashedName: step-4
---

# --description--

Sasa umeongeza wito wa kitendakazi unaweza kuona kosa lolote lililotokea ndani ya kitendakazi kwenye terminal, na unaweza kila wakati kutumia `print` kuona thamani ya vitu.

Kuna nambari ya siri iliyofichwa katika kila mstari, kwa hivyo ndani ya kitendakazi tumia njia ya `split` kugawanya mfuatano wa herufi kuwa orodha ya mistari. Gawanya mistari kwa kutumia tabia ya newline (`\n`), na weka orodha iliyopatikana katika kigezo kinachoitwa `lines`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `lines` ndani ya kitendakazi chako cha `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('lines')
    `)
})
```

Kigezo `lines` kinapaswa kuwa na thamani ya `poem.split('\n')`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('lines').is_equivalent('lines = poem.split("\\\\n")')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
```
