---
id: 68553423d775cc60134ac8bb
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Unahitaji kufanya kazi kwa kila mstari kwa kujitegemea: tengeneza mzunguko juu ya `lines` unaotumia `line` kama kigezo cha mzunguko. Ndani ya mzunguko, chapisha `line`.

# --hints--

Unapaswa kuwa na mzunguko wa `for` katika kitendakazi chako cha `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0]
    `)
})
```

Mzunguko unapaswa kurudia juu ya `lines`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('lines')
    `)
})
```

Mzunguko unapaswa kuwa na `line` kama kigezo cha mzunguko (`for line in lines`).

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line')
    `)
})
```

Unapaswa kuwa na `print(line)` ndani ya mzunguko.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
