---
id: 6855918e87c33c6979d21d3f
title: Hatua 14
challengeType: 20
dashedName: step-14
---

# --description--

Sasa kitendakazi kinafanya kazi nzuri ya kutoa msimbo kutoka kwenye shairi, lakini kuna kesi maalum ya kuzingatia: mstari wa shairi unaweza kuwa mfupi kuliko ilivyotarajiwa, na hiyo itasababisha kitendakazi kutoa kosa (unaweza kujaribu hii kwa kuondoa neno `bright` kutoka kwenye shairi).

Weka mstari `secret_code += str(len(words[line_index]))` katika taarifa ya `if` inayothibitisha kuwa kuna maneno ya kutosha katika orodha ya `words`.

# --hints--

Ndani ya mzunguko unapaswa kuwa na taarifa ya `if`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0]
    `)
})
```

Masharti ya taarifa ya if yanapaswa kuwa `len(words) > line_index`.


```js
({
    test: () => runPython(`
    _cond = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
    assert _cond.is_equivalent('len(words) > line_index') or _cond.is_equivalent('line_index < len(words)')
    `)
})
```

Ndani ya taarifa ya `if` unapaswa kuwa na `secret_code += str(len(words[line_index]))`.


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt('secret_code += str(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
--fcc-editable-region--
        secret_code += str(len(words[line_index]))
--fcc-editable-region--
    return secret_code

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

print(pin_extractor(poem))

```
