---
id: 685535bf527b2c77236b3a9a
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Nambari ya n ya siri imefichwa kama urefu wa neno la n katika mstari wa n.

Ili kupata urefu wa neno la n, hatua inayofuata ni kugawanya mstari wa shairi kuwa orodha ya maneno: ndani ya mzunguko, tengeneza kigezo `words` na weka thamani yake kuwa `line` iliyogawanywa kwa maneno kwa kutumia njia ya `split`.

Kisha, ongeza wito wa `print` ukiwa na `words` kama hoja yake.

# --hints--

Unapaswa kuwa na kigezo cha `words` ndani ya mzunguko.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_variable('words')
    `)
})
```

Kigezo cha `words` kinapaswa kuwa na thamani ya `line.split()`.

```js
({
    test: () => runPython(`
    variable = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_variable('words')
    assert variable.is_equivalent('words = line.split()') or variable.is_equivalent('words = line.split(" ")')
    `)
})
```

Unapaswa kuwa na `print(words)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        print(line)
--fcc-editable-region--
        
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
