---
id: 685578aae65e4106ad2b4300
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Kama ulivyojifunza katika funzo lililopita, kitendakazi cha `enumerate` kinakuruhusu kufuatilia kielezo wakati unaporudia juu ya inayoweza kurudiwa:

```py
fruits = ['apple', 'plum', 'bananas']

for index, item in enumerate(fruits):
    print(index, item)

# 0 apple
# 1 plum
# 2 bananas
```

Nambari ya n ya msimbo wa siri hutokana na neno la n la mstari wa n, kwa hivyo unahitaji kujua ni nambari gani ya kila mstari.

Badilisha mzunguko ili urudie juu ya `enumerate(lines)`, na ongeza kigezo kingine cha mzunguko kabla ya `line` kinachoitwa `line_index`.

Pia, sasisha wito wa `print` kuwa `print(line_index, line)`.

# --hints--

Mzunguko unapaswa kurudia juu ya `enumerate(lines)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('enumerate(lines)')
    `)
})
```

Mzunguko unapaswa kuwa na `line_index, line` kama kigezo cha mzunguko.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line_index, line')
    `)
})
```

Unapaswa kuwa na `print(line_index, line)` ndani ya mzunguko.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line_index, line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    for line in lines:
        print(line)
--fcc-editable-region--
        words = line.split()
        print(words)

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
