---
id: 68559f3561d7df287586142e
title: Hatua 16
challengeType: 20
dashedName: step-16
---

# --description--

Sasa kivunja pini kinafanya kazi, lakini kinaweza kuwa na ufanisi zaidi! Kinaweza kutoa pini kutoka kwa mashairi mengi kwa wakati mmoja!

Ongeza mashairi mawili zaidi katika wigo wa kimataifa, `poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'` na `poem3 = 'There\nonce\nwas\na\ndragon'`, na kwa sasa, toa maoni kwenye `print` yenye wito wa `pin_extractor`.


# --hints--

Vigezo `poem2` na `poem3` vinapaswa kuwepo na kuwa na thamani sahihi.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('poem2')
        assert _Node(_code).has_variable('poem3')
        assert _Node(_code).find_variable('poem2').is_equivalent("poem2 = 'The grass is green\\\\nhere and there\\\\nhoping for rain\\\\nbefore it turns yellow'")
        assert _Node(_code).find_variable('poem3').is_equivalent("poem3 = 'There\\\\nonce\\\\nwas\\\\na\\\\ndragon'")
    `)
})
```

Unapaswa kutoa maoni kwenye wito wa `print(pin_extractor(poem))`.

```js
({
    test: () => runPython(`
        assert not _Node(_code).has_call('print(pin_extractor(poem))')`)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
        else:
            secret_code += '0'
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

--fcc-editable-region--


print(pin_extractor(poem))
--fcc-editable-region--

```
