---
id: 68592730aaf34b61ae279168
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Urefu unaorejeshwa kutoka `len` ni nambari kamili, unahitaji kuibadilisha kuwa mfuatano wa herufi kwa kutumia kitendakazi cha `str` ili uweze kuiongeza kwenye mfuatano wa herufi wa `secret_code`.

Sahihisha wito wa `print` ili uchapishe urefu kama mfuatano wa herufi.

# --hints--

Unapaswa kuwa na `print(str(len(words[line_index])))` mwishoni mwa mzunguko.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(str(len(words[line_index])))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
--fcc-editable-region--
        print(len(words[line_index]))
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
