---
id: 68592b42b9c7f27cbe123203
title: Hatua 15
challengeType: 20
dashedName: step-15
---

# --description--

Mstari wa tatu wa shairi unakosa neno la tatu, kwa hivyo pini ni fupi kuliko ilivyotarajiwa.

Ongeza kifungu cha `else` kinachounganisha `'0'` na `secret_code` wakati maneno hayatoshi, ili mistari yote ya shairi itumike kupata nambari.

# --hints--

Unapaswa kuwa na kifungu cha `else`.


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1] == _Node()
    `)
})
```

Kifungu cha `else` kinapaswa kuwa na `secret_code += '0'`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_if('len(words) > line_index').is_equivalent('if len(words) > line_index:\\n  secret_code += str(len(words[line_index]))\\nelse:\\n  secret_code += \\'0\\'')
    `)
})
```

`pin_extractor(poem)` inapaswa kurudisha `5202`.


```js
({
    test: () => runPython(`
    assert pin_extractor(poem) == '5202'
    `)
})
```

`secret_code('There\nonce\nwas\na\ndragon')` inapaswa kurudisha `50000`.

```js
({
    test: () => runPython(`
    assert pin_extractor('There\\nonce\\nwas\\na\\ndragon') == '50000'
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
--fcc-editable-region--
        
--fcc-editable-region--
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

print(pin_extractor(poem))

```
