---
id: 587d7dac367417b2b2512b74
title: 점 표기법을 사용해 객체의 속성에 접근하기
challengeType: 1
forumTopicId: 301333
dashedName: use-dot-notation-to-access-the-properties-of-an-object
---

# --description--

이전 도전 과제에서는 다양한 속성을 가진 객체를 만들었습니다. 이제 그 속성들의 값을 어떻게 접근하는지 살펴보겠습니다. 예시는 다음과 같습니다:

```js
let duck = {
  name: "Aflac",
  numLegs: 2
};
console.log(duck.name);
```

점 표기법은 객체 이름 `duck` 뒤에 속성 이름 `name`을 붙여 `Aflac`의 값을 접근하는 데 사용합니다.

# --instructions--

`dog` 객체의 두 속성을 모두 콘솔에 출력하세요.

# --hints--

코드는 `console.log`를 사용해 `name` 객체의 `dog` 속성 값을 출력해야 합니다.

```js
assert(/console.log\(.*dog\.name.*\)/g.test(__helpers.removeJSComments(code)));
```

코드는 `console.log`를 사용해 `numLegs` 객체의 `dog` 속성 값을 출력해야 합니다.

```js
assert(/console.log\(.*dog\.numLegs.*\)/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4
};
// Only change code below this line
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4
};
console.log(dog.name);
console.log(dog.numLegs);
```
