---
id: 587d7dae367417b2b2512b79
title: 생성자가 인수를 받도록 확장하기
challengeType: 1
forumTopicId: 18235
dashedName: extend-constructors-to-receive-arguments
---

# --description--

지난 도전 과제의 `Bird`와 `Dog` 생성자는 잘 작동했습니다. 하지만 `Birds` 생성자로 만들어진 모든 `Bird`는 자동으로 이름이 Albert이고, 색상은 파란색이며, 다리가 두 개라는 점을 주목하세요. 만약 이름과 색상이 다른 새를 원한다면 어떻게 해야 할까요? 각 새의 속성을 수동으로 바꾸는 것도 가능하지만, 그건 매우 많은 작업이 될 것입니다:

```js
let swan = new Bird();
swan.name = "Carlos";
swan.color = "white";
```

수백 또는 수천 마리의 서로 다른 새를 조류관에서 관리하는 프로그램을 작성한다고 가정해 보세요. 모든 새를 만들고 각각의 속성을 다른 값으로 바꾸는 데 많은 시간이 걸릴 것입니다. 서로 다른 `Bird` 객체를 더 쉽게 만들려면 Bird 생성자가 매개변수를 받도록 설계할 수 있습니다:

```js
function Bird(name, color) {
  this.name = name;
  this.color = color;
  this.numLegs = 2;
}
```

그런 다음 `Bird` 생성자에 인수로 값을 전달하여 각 고유한 새를 정의하세요: `let cardinal = new Bird("Bruce", "red");` 이렇게 하면 `Bird`의 새 인스턴스가 생성되며 `name`와 `color` 속성이 각각 `Bruce`와 `red`로 설정됩니다. `numLegs` 속성은 여전히 2로 설정되어 있습니다. `cardinal`는 다음과 같은 속성을 가집니다:

```js
cardinal.name
cardinal.color
cardinal.numLegs
```

생성자가 더 유연해졌습니다. 이제 각 `Bird`가 생성될 때 속성을 정의할 수 있는데, 이것이 JavaScript 생성자가 매우 유용한 한 가지 이유입니다. 생성자는 공통된 특성과 동작을 기준으로 객체를 그룹화하고, 그들의 생성을 자동화하는 청사진을 정의합니다.

# --instructions--

또 다른 `Dog` 생성자를 만드세요. 이번에는 `name`와 `color` 매개변수를 받도록 설정하고, `numLegs` 속성은 4로 고정하세요. 그런 다음 `Dog`를 새로 만들어 `terrier` 변수에 저장하세요. `name`와 `color` 속성에 대해 두 개의 문자열을 인수로 전달하세요.

# --hints--

`Dog`는 `name`에 대한 인수를 받아야 합니다.

```js
assert(new Dog('Clifford').name === 'Clifford');
```

`Dog`는 `color`에 대한 인수를 받아야 합니다.

```js
assert(new Dog('Clifford', 'yellow').color === 'yellow');
```

`Dog`는 `numLegs` 속성을 4로 설정해야 합니다.

```js
assert(new Dog('Clifford').numLegs === 4);
```

`terrier`는 `Dog` 생성자를 사용해 만들어야 합니다.

```js
assert(terrier instanceof Dog);
```

# --seed--

## --seed-contents--

```js
function Dog() {

}
```

# --solutions--

```js
function Dog (name, color) {
  this.numLegs = 4;
  this.name = name;
  this.color = color;
}

const terrier = new Dog();
```
