---
id: 587d7dad367417b2b2512b76
title: this 키워드로 코드를 더 재사용 가능하게 만들기
challengeType: 1
forumTopicId: 301321
dashedName: make-code-more-reusable-with-the-this-keyword
---

# --description--

마지막 도전 과제에서는 `duck` 객체에 메서드를 추가했습니다. `duck.name` 점 표기법을 사용해 return 문 안에서 `name` 속성의 값을 접근했습니다:

```js
sayName: function() {return "The name of this duck is " + duck.name + ".";}
```

이 방법은 객체의 속성에 접근하는 유효한 방법이지만, 여기에는 함정이 있습니다. 변수 이름이 바뀌면 원래 이름을 참조하는 모든 코드를 업데이트해야 합니다. 짧은 객체 정의에서는 문제가 되지 않지만, 객체가 많은 속성 참조를 포함하면 오류가 발생할 가능성이 커집니다.

이 문제를 피하는 방법은 `this` 키워드를 사용하는 것입니다:

```js
let duck = {
  name: "Aflac",
  numLegs: 2,
  sayName: function() {return "The name of this duck is " + this.name + ".";}
};
```

`this`는 깊이 있는 주제이며, 위 예시는 그중 한 가지 사용법일 뿐입니다. 현재 문맥에서 `this`는 메서드가 연결된 객체인 `duck`를 가리킵니다. 객체 이름이 `mallard`로 바뀌어도 코드 내 `duck`에 대한 모든 참조를 찾을 필요가 없습니다. 이렇게 하면 코드가 재사용 가능하고 읽기 쉬워집니다.

# --instructions--

`dog.sayLegs` 메서드를 수정해 `dog`에 대한 모든 참조를 제거하세요. `duck` 예시를 참고하세요.

# --hints--

`dog.sayLegs()`는 주어진 문자열을 반환해야 합니다.

```js
assert(dog.sayLegs() === 'This dog has 4 legs.');
```

코드는 `this` 키워드를 사용해 `numLegs`의 `dog` 속성에 접근해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/this\.numLegs/g));
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4,
  sayLegs: function() {return "This dog has " + dog.numLegs + " legs.";}
};

dog.sayLegs();
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4,
  sayLegs () {
    return 'This dog has ' + this.numLegs + ' legs.';
  }
};

dog.sayLegs();
```
