---
id: 587d7db1367417b2b2512b86
title: 상속된 생성자 속성 재설정
challengeType: 1
forumTopicId: 301324
dashedName: reset-an-inherited-constructor-property
---

# --description--

객체가 다른 객체로부터 `prototype`를 상속받으면, 슈퍼타입의 생성자 속성도 함께 상속받습니다.

예를 들어 보겠습니다:

```js
function Bird() { }
Bird.prototype = Object.create(Animal.prototype);
let duck = new Bird();
duck.constructor
```

하지만 `duck`와 `Bird`의 모든 인스턴스는 `Bird`가 아니라 `Animal`에 의해 생성되었음을 보여야 합니다. 이를 위해 `Bird`의 생성자 속성을 `Bird` 객체로 수동 설정할 수 있습니다:

```js
Bird.prototype.constructor = Bird;
duck.constructor
```

# --instructions--

`duck.constructor`와 `beagle.constructor`가 각각 자신의 생성자를 반환하도록 코드를 수정하세요.

# --hints--

`Bird.prototype`는 `Animal`의 인스턴스여야 합니다.

```js
assert(Animal.prototype.isPrototypeOf(Bird.prototype));
```

`duck.constructor`는 `Bird`를 반환해야 합니다.

```js
assert(duck.constructor === Bird);
```

`Dog.prototype`는 `Animal`의 인스턴스여야 합니다.

```js
assert(Animal.prototype.isPrototypeOf(Dog.prototype));
```

`beagle.constructor`는 `Dog`를 반환해야 합니다.

```js
assert(beagle.constructor === Dog);
```

# --seed--

## --seed-contents--

```js
function Animal() { }
function Bird() { }
function Dog() { }

Bird.prototype = Object.create(Animal.prototype);
Dog.prototype = Object.create(Animal.prototype);

// Only change code below this line



let duck = new Bird();
let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }
function Bird() { }
function Dog() { }
Bird.prototype = Object.create(Animal.prototype);
Dog.prototype = Object.create(Animal.prototype);
Dog.prototype.constructor = Dog;
Bird.prototype.constructor = Bird;
let duck = new Bird();
let beagle = new Dog();
```
