---
id: 587d7db2367417b2b2512b8a
title: >-
  클로저를 사용해 객체 내부 속성이 외부에서 수정되지 않도록 보호하기
challengeType: 1
forumTopicId: 18234
dashedName: >-
  use-closure-to-protect-properties-within-an-object-from-being-modified-externally
---

# --description--

이전 도전 과제에서 `bird`는 공개 속성 `name`을 가지고 있었습니다. `bird` 정의 외부에서 접근하고 변경할 수 있기 때문에 공개된 것으로 간주됩니다.

```js
bird.name = "Duffy";
```

따라서 코드의 어느 부분에서든 `bird`의 이름을 어떤 값으로든 쉽게 변경할 수 있습니다. 비밀번호나 은행 계좌 같은 것이 코드베이스의 어느 부분에서든 쉽게 변경될 수 있다고 생각해 보세요. 이는 많은 문제를 일으킬 수 있습니다.

이 공개 속성을 비공개로 만드는 가장 간단한 방법은 생성자 함수 내에 변수를 만드는 것입니다. 이렇게 하면 그 변수의 범위가 전역이 아닌 생성자 함수 내부로 변경됩니다. 따라서 그 변수는 생성자 함수 내의 메서드에서만 접근하고 변경할 수 있습니다.

```js
function Bird() {
  let hatchedEgg = 10;

  this.getHatchedEggCount = function() { 
    return hatchedEgg;
  };
}
let ducky = new Bird();
ducky.getHatchedEggCount();
```

여기서 `getHatchedEggCount`는 특권 메서드인데, 이는 비공개 변수 `hatchedEgg`에 접근할 수 있기 때문입니다. `hatchedEgg`가 `getHatchedEggCount`과 같은 범위에서 선언되었기 때문에 가능합니다. JavaScript에서 함수는 항상 자신이 생성된 범위에 접근할 수 있습니다. 이것을 `closure`라고 합니다.

# --instructions--

`weight` 함수 내에서 `Bird`가 선언되는 방식을 변경해 비공개 변수가 되도록 하세요. 그리고 `getWeight` 15의 값을 반환하는 `weight` 메서드를 만드세요.

# --hints--

`weight` 속성은 비공개 변수여야 하며 `15` 값을 할당받아야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/(var|let|const)\s+weight\s*\=\s*15\;?/g));
```

코드는 `Bird`에 `getWeight`라는 메서드를 만들어 비공개 변수 `weight`의 값을 반환해야 합니다.

```js
assert(new Bird().getWeight() === 15);
```

`getWeight` 함수는 비공개 변수 `weight`를 반환해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/((return\s+)|(\(\s*\)\s*\=\>\s*))weight\;?/g));
```

# --seed--

## --seed-contents--

```js
function Bird() {
  this.weight = 15;


}
```

# --solutions--

```js
function Bird() {
  let weight = 15;

  this.getWeight = () => weight;
}
```
