---
id: 587d7db1367417b2b2512b87
title: 상속 후 메서드 추가
challengeType: 1
forumTopicId: 301315
dashedName: add-methods-after-inheritance
---

# --description--

슈퍼타입 생성자 함수로부터 `prototype` 객체를 상속받는 생성자 함수는 상속받은 메서드 외에 자체 메서드를 가질 수 있습니다.

예를 들어, `Bird`은 `prototype`로부터 `Animal`를 상속받는 생성자입니다:

```js
function Animal() { }
Animal.prototype.eat = function() {
  console.log("nom nom nom");
};
function Bird() { }
Bird.prototype = Object.create(Animal.prototype);
Bird.prototype.constructor = Bird;
```

`Animal`로부터 상속받은 것 외에 `Bird` 객체에만 고유한 동작을 추가하고 싶습니다. 여기서 `Bird`는 `fly()` 함수를 갖게 됩니다. 함수는 다른 생성자 함수와 같은 방식으로 `Bird's` `prototype`에 추가됩니다:

```js
Bird.prototype.fly = function() {
  console.log("I'm flying!");
};
```

이제 `Bird`의 인스턴스는 `eat()`과 `fly()` 메서드를 모두 갖게 됩니다:

```js
let duck = new Bird();
duck.eat();
duck.fly();
```

`duck.eat()`은 콘솔에 문자열 `nom nom nom`를 출력하고, `duck.fly()`는 문자열 `I'm flying!`를 출력할 것입니다.

# --instructions--

필요한 모든 코드를 추가하여 `Dog` 객체가 `Animal`를 상속받고 `Dog`의 `prototype` 생성자가 `Dog`로 설정되게 하세요. 그런 다음 `bark()` 객체에 `Dog` 메서드를 추가하여 `beagle`가 `eat()`와 `bark()`를 모두 할 수 있게 하세요. `bark()` 메서드는 콘솔에 `Woof!`를 출력해야 합니다.

# --hints--

`Animal`는 `bark()` 메서드에 반응하지 않아야 합니다.

```js
assert(typeof Animal.prototype.bark == 'undefined');
```

`Dog`는 `eat()`로부터 `Animal` 메서드를 상속받아야 합니다.

```js
assert(typeof Dog.prototype.eat == 'function');
```

`Dog` 프로토타입은 `bark()` 메서드를 가져야 합니다.

```js
assert('bark' in Dog.prototype);
```

`beagle`는 `instanceof` `Animal`이어야 합니다.

```js
assert(beagle instanceof Animal);
```

`beagle`의 생성자는 `Dog`로 설정되어야 합니다.

```js
assert(beagle.constructor === Dog);
```

`beagle.eat()`는 문자열 `nom nom nom`를 출력해야 합니다.

```js
const spy = __helpers.spyOn(console, 'log');
beagle.eat();
assert.lengthOf(spy.calls,1);
assert(spy.calls[0][0] == 'nom nom nom');
```

`beagle.bark()`는 문자열 `Woof!`를 출력해야 합니다.

```js
const spy = __helpers.spyOn(console, 'log');
beagle.bark();
assert.lengthOf(spy.calls,1);
assert(spy.calls[0][0] == 'Woof!');
```

# --seed--

## --seed-contents--

```js
function Animal() { }
Animal.prototype.eat = function() { console.log("nom nom nom"); };

function Dog() { }

// Only change code below this line




// Only change code above this line

let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }
Animal.prototype.eat = function() { console.log("nom nom nom"); };

function Dog() { }
Dog.prototype = Object.create(Animal.prototype);
Dog.prototype.constructor = Dog;
Dog.prototype.bark = function () {
  console.log('Woof!');
};
let beagle = new Dog();

beagle.eat();
beagle.bark();
```
