---
id: 587d7daf367417b2b2512b7d
title: 모든 프로퍼티 순회하기
challengeType: 1
forumTopicId: 301320
dashedName: iterate-over-all-properties
---

# --description--

지금까지 두 가지 종류의 프로퍼티를 보았습니다: <dfn>자체 프로퍼티</dfn>와 `prototype` 프로퍼티입니다. 자체 프로퍼티는 객체 인스턴스 자체에 직접 정의됩니다. 그리고 프로토타입 프로퍼티는 `prototype`에 정의됩니다.

```js
function Bird(name) {
  this.name = name;  //own property
}

Bird.prototype.numLegs = 2; // prototype property

let duck = new Bird("Donald");
```

다음은 `duck`의 자체 프로퍼티를 배열 `ownProps`에, `prototype` 프로퍼티를 배열 `prototypeProps`에 추가하는 방법입니다:

```js
let ownProps = [];
let prototypeProps = [];

for (let property in duck) {
  if(duck.hasOwnProperty(property)) {
    ownProps.push(property);
  } else {
    prototypeProps.push(property);
  }
}

console.log(ownProps);
console.log(prototypeProps);
```

`console.log(ownProps)`는 콘솔에 `["name"]`를 표시하고, `console.log(prototypeProps)`는 `["numLegs"]`를 표시합니다.

# --instructions--

`beagle`의 모든 자체 프로퍼티를 배열 `ownProps`에 추가하세요. `prototype`의 모든 `Dog` 프로퍼티를 배열 `prototypeProps`에 추가하세요.

# --hints--

`ownProps` 배열에는 `name`만 포함되어야 합니다.

```js
assert.deepEqual(ownProps, ['name']);
```

`prototypeProps` 배열에는 `numLegs`만 포함되어야 합니다.

```js
assert.deepEqual(prototypeProps, ['numLegs']);
```

내장 메서드 `Object.keys()`를 사용하지 않고 이 도전 과제를 해결하세요.

```js
assert(!/\Object.keys/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype.numLegs = 4;

let beagle = new Dog("Snoopy");

let ownProps = [];
let prototypeProps = [];

// Only change code below this line
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype.numLegs = 4;

let beagle = new Dog("Snoopy");

let ownProps = [];
let prototypeProps = [];
for (let prop in beagle) {
  if (beagle.hasOwnProperty(prop)) {
    ownProps.push(prop);
  } else {
    prototypeProps.push(prop);
  }
}
```
