---
id: 587d7dad367417b2b2512b75
title: 객체에 메서드 만들기
challengeType: 1
forumTopicId: 301318
dashedName: create-a-method-on-an-object
---

# --description--

객체는 <dfn>메서드</dfn>라고 하는 특별한 유형의 속성을 가질 수 있습니다.

메서드는 함수인 속성입니다. 이것은 객체에 다른 동작을 추가합니다. 다음은 메서드가 포함된 `duck` 예제입니다:

```js
let duck = {
  name: "Aflac",
  numLegs: 2,
  sayName: function() {return "The name of this duck is " + duck.name + ".";}
};
duck.sayName();
```

이 예제는 `sayName` 메서드를 추가하는데, 이 메서드는 `duck`의 이름을 반환하는 문장을 반환하는 함수입니다. 메서드가 `name`를 사용해 return 문에서 `duck.name` 속성에 접근한 점을 주목하세요. 다음 도전 과제에서는 이를 수행하는 또 다른 방법을 다룹니다.

# --instructions--

`dog` 객체를 사용해 `sayLegs`라는 메서드를 만드세요. 이 메서드는 `This dog has 4 legs.` 문장을 반환해야 합니다.

# --hints--

`dog.sayLegs()`는 함수여야 합니다.

```js
assert(typeof dog.sayLegs === 'function');
```

`dog.sayLegs()`는 주어진 문자열을 반환해야 합니다 — 구두점과 띄어쓰기에 주의하세요.

```js
assert(dog.sayLegs() === 'This dog has 4 legs.');
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4,

};

dog.sayLegs();
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4,
  sayLegs () {
    return 'This dog has ' + this.numLegs + ' legs.';
  }
};

dog.sayLegs();
```
