---
id: 587d7dad367417b2b2512b78
title: 생성자를 사용해 객체 만들기
challengeType: 1
forumTopicId: 18233
dashedName: use-a-constructor-to-create-objects
---

# --description--

이전 도전 과제에서 사용한 `Bird` 생성자는 다음과 같습니다:

```js
function Bird() {
  this.name = "Albert";
  this.color  = "blue";
  this.numLegs = 2;
}

let blueBird = new Bird();
```

**참고:** 생성자 내부의 `this`는 항상 생성 중인 객체를 가리킵니다.

생성자를 호출할 때 `new` 연산자가 사용된다는 점에 주목하세요. 이 연산자는 JavaScript에 `Bird`의 새 인스턴스인 `blueBird`를 생성하라고 알려줍니다. `new` 연산자가 없으면 생성자 내부의 `this`는 새로 생성된 객체를 가리키지 않아 예상치 못한 결과가 발생합니다. 이제 `blueBird`는 `Bird` 생성자 내부에 정의된 모든 프로퍼티를 갖게 됩니다:

```js
blueBird.name;
blueBird.color;
blueBird.numLegs;
```

다른 객체와 마찬가지로, 그 프로퍼티들은 접근하고 수정할 수 있습니다:

```js
blueBird.name = 'Elvira';
blueBird.name;
```

# --instructions--

지난 수업에서 사용한 `Dog` 생성자를 사용해 `Dog`의 새 인스턴스를 만들고, 이를 `hound` 변수에 할당하세요.

# --hints--

`hound`는 `Dog` 생성자를 사용해 만들어야 합니다.

```js
assert(hound instanceof Dog);
```

코드는 `new` 연산자를 사용해 `Dog`의 인스턴스를 생성해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/new/g));
```

# --seed--

## --seed-contents--

```js
function Dog() {
  this.name = "Rupert";
  this.color = "brown";
  this.numLegs = 4;
}
// Only change code below this line
```

# --solutions--

```js
function Dog() {
  this.name = "Rupert";
  this.color = "brown";
  this.numLegs = 4;
}
const hound = new Dog();
```
