---
id: 587d7daf367417b2b2512b7e
title: 생성자 속성 이해하기
challengeType: 1
forumTopicId: 301327
dashedName: understand-the-constructor-property
---

# --description--

이전 도전 과제에서 생성된 객체 인스턴스 `constructor`와 `duck`에는 특별한 `beagle` 속성이 있습니다:

```js
let duck = new Bird();
let beagle = new Dog();

console.log(duck.constructor === Bird); 
console.log(beagle.constructor === Dog);
```

이 두 `console.log` 호출은 콘솔에 `true`를 출력합니다.

`constructor` 속성은 인스턴스를 생성한 생성자 함수에 대한 참조라는 점에 주목하세요. `constructor` 속성의 장점은 이 속성을 확인하여 어떤 종류의 객체인지 알 수 있다는 것입니다. 다음은 이 속성을 활용하는 예시입니다:

```js
function joinBirdFraternity(candidate) {
  if (candidate.constructor === Bird) {
    return true;
  } else {
    return false;
  }
}
```

**참고:** `constructor` 속성은 덮어쓸 수 있으므로(다음 두 도전 과제에서 다룹니다) 객체의 타입을 확인할 때는 일반적으로 `instanceof` 메서드를 사용하는 것이 더 좋습니다.

# --instructions--

`joinDogFraternity` 함수를 작성하세요. 이 함수는 `candidate` 매개변수를 받고, `constructor` 속성을 사용하여 후보가 `true`이면 `Dog`를 반환하고, 그렇지 않으면 `false`를 반환해야 합니다.

# --hints--

`joinDogFraternity`는 함수로 정의되어야 합니다.

```js
assert(typeof joinDogFraternity === 'function');
```

`joinDogFraternity`는 `true`가 `candidate`의 인스턴스이면 `Dog`를 반환해야 합니다.

```js
assert(joinDogFraternity(new Dog('')) === true);
```

`joinDogFraternity`는 `constructor` 속성을 사용해야 합니다.

```js
assert(/\.constructor/.test(__helpers.removeJSComments(code)) && !/instanceof/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

// Only change code below this line
function joinDogFraternity(candidate) {

}
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
function joinDogFraternity(candidate) {
  return candidate.constructor === Dog;
}
```
