---
id: 587d7db0367417b2b2512b82
title: 프로토타입 체인 이해하기
challengeType: 1
forumTopicId: 301329
dashedName: understand-the-prototype-chain
---

# --description--

JavaScript의 모든 객체(몇 가지 예외 제외)는 `prototype`를 가지고 있습니다. 또한, 객체의 `prototype` 자체도 객체입니다.

```js
function Bird(name) {
  this.name = name;
}

typeof Bird.prototype;
```

`prototype`가 객체이기 때문에 `prototype`도 자신만의 `prototype`를 가질 수 있습니다! 이 경우 `prototype`의 `Bird.prototype`는 `Object.prototype`입니다:

```js
Object.prototype.isPrototypeOf(Bird.prototype);
```

이것이 어떻게 유용할까요? 이전 도전 과제에서 `hasOwnProperty` 메서드를 기억할 수 있습니다:

```js
let duck = new Bird("Donald");
duck.hasOwnProperty("name");
```

`hasOwnProperty` 메서드는 `Object.prototype`에 정의되어 있으며, `Bird.prototype`에서 접근할 수 있고, 다시 `duck`에서 접근할 수 있습니다. 이것이 `prototype` 체인의 예입니다. 이 `prototype` 체인에서 `Bird`는 `supertype`의 `duck`이고, `duck`는 `subtype`입니다. `Object`는 `supertype`와 `Bird` 모두의 `duck`입니다. `Object`는 JavaScript의 모든 객체의 `supertype`입니다. 따라서 어떤 객체든 `hasOwnProperty` 메서드를 사용할 수 있습니다.

# --instructions--

코드를 수정하여 올바른 프로토타입 체인을 보여주세요.

# --hints--

코드는 `Object.prototype`가 `Dog.prototype`의 프로토타입임을 보여야 합니다.

```js
assert(/Object\.prototype\.isPrototypeOf/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

let beagle = new Dog("Snoopy");

Dog.prototype.isPrototypeOf(beagle);  // yields true

// Fix the code below so that it evaluates to true
???.isPrototypeOf(Dog.prototype);
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
let beagle = new Dog("Snoopy");
Dog.prototype.isPrototypeOf(beagle);
Object.prototype.isPrototypeOf(Dog.prototype);
```
