---
id: 587d7dae367417b2b2512b7a
title: instanceof로 객체의 생성자 확인하기
challengeType: 1
forumTopicId: 301337
dashedName: verify-an-objects-constructor-with-instanceof
---

# --description--

생성자 함수가 새 객체를 만들 때마다, 그 객체는 해당 생성자의 <dfn>인스턴스</dfn>라고 합니다. JavaScript는 `instanceof` 연산자를 사용해 이를 편리하게 확인할 수 있습니다. `instanceof`는 객체와 생성자를 비교해, 그 객체가 해당 생성자로 만들어졌는지에 따라 `true` 또는 `false`를 반환합니다. 예시는 다음과 같습니다:

```js
let Bird = function(name, color) {
  this.name = name;
  this.color = color;
  this.numLegs = 2;
}

let crow = new Bird("Alexis", "black");

crow instanceof Bird;
```

이 `instanceof` 메서드는 `true`를 반환합니다.

객체가 생성자를 사용하지 않고 만들어졌다면, `instanceof`는 그 객체가 해당 생성자의 인스턴스가 아님을 확인합니다:

```js
let canary = {
  name: "Mildred",
  color: "Yellow",
  numLegs: 2
};

canary instanceof Bird;
```

이 `instanceof` 메서드는 `false`를 반환합니다.

# --instructions--

`House` 생성자의 새 인스턴스를 만들고, `myHouse`이라고 부르며 침실 수를 전달하세요. 그런 다음 `instanceof`를 사용해 `House`의 인스턴스인지 확인하세요.

# --hints--

`myHouse`는 숫자로 설정된 `numBedrooms` 속성을 가져야 합니다.

```js
assert(typeof myHouse.numBedrooms === 'number');
```

`myHouse`가 `House`의 인스턴스인지 `instanceof` 연산자를 사용해 확인하세요.

```js
assert(/myHouse\s*instanceof\s*House/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function House(numBedrooms) {
  this.numBedrooms = numBedrooms;
}

// Only change code below this line
```

# --solutions--

```js
function House(numBedrooms) {
  this.numBedrooms = numBedrooms;
}
const myHouse = new House(4);
console.log(myHouse instanceof House);
```
