---
id: 587d7daf367417b2b2512b80
title: 프로토타입을 변경할 때 생성자 속성을 설정하는 것을 잊지 마세요
challengeType: 1
forumTopicId: 301323
dashedName: remember-to-set-the-constructor-property-when-changing-the-prototype
---

# --description--

프로토타입을 새 객체로 수동 설정할 때 중요한 부작용이 하나 있습니다. `constructor` 속성이 지워집니다! 이 속성은 어떤 생성자 함수가 인스턴스를 만들었는지 확인하는 데 사용되지만, 속성이 덮어써졌기 때문에 이제 잘못된 결과를 반환합니다:

```js
duck.constructor === Bird;
duck.constructor === Object;
duck instanceof Bird;
```

이 표현식들은 차례대로 `false`, `true`, `true`로 평가됩니다.

이 문제를 고치려면, 프로토타입을 새 객체로 수동 설정할 때마다 `constructor` 속성을 정의하는 것을 기억하세요:

```js
Bird.prototype = {
  constructor: Bird,
  numLegs: 2,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name); 
  }
};
```

# --instructions--

`constructor` `Dog`에 `prototype` 속성을 정의하세요.

# --hints--

`Dog.prototype`는 `constructor` 속성을 설정해야 합니다.

```js
assert(Dog.prototype.constructor === Dog);
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

// Only change code below this line
Dog.prototype = {

  numLegs: 4,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
Dog.prototype = {
  constructor: Dog,
  numLegs: 4,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```
