---
id: 587d7db1367417b2b2512b88
title: 상속받은 메서드 재정의하기
challengeType: 1
forumTopicId: 301322
dashedName: override-inherited-methods
---

# --description--

이전 수업에서, 객체가 `prototype` 객체를 참조하여 다른 객체로부터 동작(메서드)을 상속받을 수 있다는 것을 배웠습니다.

```js
ChildObject.prototype = Object.create(ParentObject.prototype);
```

그다음 `ChildObject`은 `prototype`에 메서드를 체이닝하여 자신만의 메서드를 받았습니다.

```js
ChildObject.prototype.methodName = function() {...};
```

상속받은 메서드를 재정의하는 것도 가능합니다. 같은 방법으로, 재정의할 메서드 이름과 동일한 이름으로 `ChildObject.prototype`에 메서드를 추가하면 됩니다. 다음은 `Bird`가 `eat()`로부터 상속받은 `Animal` 메서드를 재정의한 예입니다.

```js
function Animal() { }
Animal.prototype.eat = function() {
  return "nom nom nom";
};
function Bird() { }

Bird.prototype = Object.create(Animal.prototype);

Bird.prototype.eat = function() {
  return "peck peck peck";
};
```

만약 `let duck = new Bird();` 인스턴스가 있고 `duck.eat()`를 호출하면, JavaScript는 `prototype`의 `duck` 체인에서 메서드를 다음과 같이 찾습니다:

1.  `duck` => 여기서 `eat()`가 정의되어 있나요? 아니요.
2.  `Bird` => 여기서 `eat()`가 정의되어 있나요? => 네. 실행하고 검색을 멈춥니다.
3.  `Animal` => `eat()`도 정의되어 있지만, JavaScript는 이 단계에 도달하기 전에 검색을 멈췄습니다.
4.  Object => JavaScript는 이 단계에 도달하기 전에 검색을 멈췄습니다.

# --instructions--

`fly()`의 `Penguin` 메서드를 재정의하여 문자열 `Alas, this is a flightless bird.`를 반환하게 하세요.

# --hints--

`penguin.fly()`는 문자열 `Alas, this is a flightless bird.`를 반환해야 합니다.

```js
assert(penguin.fly() === 'Alas, this is a flightless bird.');
```

`bird.fly()` 메서드는 문자열 `I am flying!`를 반환해야 합니다.

```js
assert(new Bird().fly() === 'I am flying!');
```

# --seed--

## --seed-contents--

```js
function Bird() { }

Bird.prototype.fly = function() { return "I am flying!"; };

function Penguin() { }
Penguin.prototype = Object.create(Bird.prototype);
Penguin.prototype.constructor = Penguin;

// Only change code below this line



// Only change code above this line

let penguin = new Penguin();
console.log(penguin.fly());
```

# --solutions--

```js
function Bird() { }

Bird.prototype.fly = function() { return "I am flying!"; };

function Penguin() { }
Penguin.prototype = Object.create(Bird.prototype);
Penguin.prototype.constructor = Penguin;
Penguin.prototype.fly = () => 'Alas, this is a flightless bird.';
let penguin = new Penguin();
console.log(penguin.fly());
```
