---
id: 587d7db0367417b2b2512b84
title: 슈퍼타입에서 동작 상속하기
challengeType: 1
forumTopicId: 301319
dashedName: inherit-behaviors-from-a-supertype
---

# --description--

이전 도전 과제에서는 모든 동물이 공유하는 동작을 정의한 `supertype`라는 `Animal`를 만들었습니다:

```js
function Animal() { }
Animal.prototype.eat = function() {
  console.log("nom nom nom");
};
```

이번 도전 과제와 다음 도전 과제에서는 `Animal`의 메서드를 `Bird`와 `Dog` 안에서 다시 정의하지 않고 재사용하는 방법을 다룹니다. 이 방법은 상속이라는 기법을 사용합니다. 이번 도전 과제는 첫 번째 단계인 `supertype` (또는 부모)의 인스턴스를 만드는 방법을 다룹니다. `Animal` 연산자를 사용해 `new`의 인스턴스를 만드는 방법은 이미 알고 있습니다:

```js
let animal = new Animal();
```

이 구문을 상속에 사용할 때 몇 가지 단점이 있는데, 이번 도전 과제 범위에는 너무 복잡합니다. 대신, 그런 단점이 없는 대안 방법이 있습니다:

```js
let animal = Object.create(Animal.prototype);
```

`Object.create(obj)`는 새 객체를 만들고, `obj`를 새 객체의 `prototype`로 설정합니다. `prototype`는 객체를 만드는 "레시피"와 같습니다. `prototype`의 `animal`를 `prototype`의 `Animal`로 설정하면, `animal` 인스턴스에 `Animal`의 다른 인스턴스와 같은 "레시피"를 주는 셈입니다.

```js
animal.eat();
animal instanceof Animal;
```

여기서 `instanceof` 메서드는 `true`를 반환합니다.

# --instructions--

`Object.create`를 사용해 `Animal`의 인스턴스 두 개를 `duck`와 `beagle`라는 이름으로 만드세요.

# --hints--

`duck` 변수가 정의되어야 합니다.

```js
assert(typeof duck !== 'undefined');
```

`beagle` 변수가 정의되어야 합니다.

```js
assert(typeof beagle !== 'undefined');
```

`duck` 변수는 `Object.create`로 초기화되어야 합니다.

```js
assert(
  /(let|const|var)\s{1,}duck\s*=\s*Object\.create\s*\(\s*Animal\.prototype\s*\)\s*/.test(
    __helpers.removeJSComments(code)
  )
);
```

`beagle` 변수는 `Object.create`로 초기화되어야 합니다.

```js
assert(
  /(let|const|var)\s{1,}beagle\s*=\s*Object\.create\s*\(\s*Animal\.prototype\s*\)\s*/.test(
    __helpers.removeJSComments(code)
  )
);
```

`duck`는 `prototype`의 `Animal`를 가져야 합니다.

```js
assert(duck instanceof Animal);
```

`beagle`는 `prototype`의 `Animal`를 가져야 합니다.

```js
assert(beagle instanceof Animal);
```

# --seed--

## --seed-contents--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

// Only change code below this line

let duck; // Change this line
let beagle; // Change this line
```

# --solutions--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};
let duck = Object.create(Animal.prototype);
let beagle = Object.create(Animal.prototype);

duck.eat();
beagle.eat();
```
