---
id: 587d7daf367417b2b2512b7f
title: 프로토타입을 새 객체로 변경하기
challengeType: 1
forumTopicId: 301316
dashedName: change-the-prototype-to-a-new-object
---

# --description--

지금까지는 `prototype`에 속성을 하나씩 추가해왔습니다:

```js
Bird.prototype.numLegs = 2;
```

몇 개 이상의 속성을 추가하면 이 방법은 번거로워집니다.

```js
Bird.prototype.eat = function() {
  console.log("nom nom nom");
}

Bird.prototype.describe = function() {
  console.log("My name is " + this.name);
}
```

더 효율적인 방법은 `prototype`를 이미 속성을 포함하는 새 객체로 설정하는 것입니다. 이렇게 하면 속성이 한꺼번에 추가됩니다:

```js
Bird.prototype = {
  numLegs: 2, 
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

# --instructions--

`numLegs`를 새 객체로 설정하여 `eat()`의 `describe()`에 `prototype` 속성과 `Dog`, `prototype` 두 메서드를 추가하세요.

# --hints--

`Dog.prototype`는 새 객체로 설정해야 합니다.

```js
assert(/Dog\.prototype\s*?=\s*?{/.test(__helpers.removeJSComments(code)));
```

`Dog.prototype`는 `numLegs` 속성을 가져야 합니다.

```js
assert(Dog.prototype.numLegs !== undefined);
```

`Dog.prototype`는 `eat()` 메서드를 가져야 합니다.

```js
assert(typeof Dog.prototype.eat === 'function');
```

`Dog.prototype`는 `describe()` 메서드를 가져야 합니다.

```js
assert(typeof Dog.prototype.describe === 'function');
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype = {
  // Only change code below this line

};
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
Dog.prototype = {
numLegs: 4,
  eat () {
    console.log('nom nom nom');
  },
  describe () {
    console.log('My name is ' + this.name);
  }
};
```
