---
id: 587d7db2367417b2b2512b8b
title: 즉시 실행 함수 표현식(IIFE) 이해하기
challengeType: 1
forumTopicId: 301328
dashedName: understand-the-immediately-invoked-function-expression-iife
---

# --description--

JavaScript에서 흔히 사용하는 패턴은 함수를 선언하자마자 실행하는 것입니다:

```js
(function () {
  console.log("Chirp, chirp!");
})();
```

이것은 즉시 실행되는 익명 함수 표현식이며, 즉시 `Chirp, chirp!`를 출력합니다.

함수에 이름이 없고 변수에 저장되지 않았다는 점에 주목하세요. 함수 표현식 끝의 두 개의 소괄호()가 함수를 즉시 실행 또는 호출하게 만듭니다. 이 패턴을 <dfn>즉시 실행 함수 표현식</dfn> 또는 <dfn>IIFE</dfn>라고 합니다.

# --instructions--

함수 `makeNest`을 익명 즉시 실행 함수 표현식(IIFE)으로 다시 작성하고 호출 부분은 제거하세요.

# --hints--

함수는 익명이어야 합니다.

```js
assert(/\((function|\(\))(=>|\(\)){?/.test(__helpers.removeJSComments(code).replace(/\s/g, '')));
```

함수를 즉시 호출하려면 표현식 끝에 소괄호가 있어야 합니다.

```js
assert(/\(.*(\)\(|\}\(\))\)/.test(__helpers.removeJSComments(code).replace(/[\s;]/g, '')));
```

# --seed--

## --seed-contents--

```js
function makeNest() {
  console.log("A cozy nest is ready");
}

makeNest();
```

# --solutions--

```js
(function () {
  console.log("A cozy nest is ready");
})();
```

---

```js
(function () {
  console.log("A cozy nest is ready");
}());
```

---

```js
(() => {
  console.log("A cozy nest is ready");
})();
```

---

```js
(() =>
  console.log("A cozy nest is ready")
)();
```
