---
id: 587d7db2367417b2b2512b8c
title: IIFE를 사용해 모듈 만들기
challengeType: 1
forumTopicId: 301332
dashedName: use-an-iife-to-create-a-module
---

# --description--

즉시 실행 함수 표현식(IIFE)은 관련 기능을 하나의 객체나 <dfn>모듈</dfn>로 묶을 때 자주 사용됩니다. 예를 들어, 이전 도전 과제에서는 두 개의 믹스인을 정의했습니다:

```js
function glideMixin(obj) {
  obj.glide = function() {
    console.log("Gliding on the water");
  };
}
function flyMixin(obj) {
  obj.fly = function() {
    console.log("Flying, wooosh!");
  };
}
```

이 믹스인들을 다음과 같이 모듈로 묶을 수 있습니다:

```js
let motionModule = (function () {
  return {
    glideMixin: function(obj) {
      obj.glide = function() {
        console.log("Gliding on the water");
      };
    },
    flyMixin: function(obj) {
      obj.fly = function() {
        console.log("Flying, wooosh!");
      };
    }
  }
})();
```

즉시 실행 함수 표현식(IIFE)이 객체 `motionModule`를 반환하는 것을 확인하세요. 이 반환된 객체는 모든 믹스인 동작을 객체의 프로퍼티로 포함합니다. 모듈 패턴의 장점은 모든 동작 동작을 하나의 객체로 패키징하여 코드의 다른 부분에서 사용할 수 있다는 점입니다. 다음은 이를 사용하는 예시입니다:

```js
motionModule.glideMixin(duck);
duck.glide();
```

# --instructions--

두 믹스인 `funModule`와 `isCuteMixin`를 감싸는 `singMixin`이라는 모듈을 만드세요. `funModule`는 객체를 반환해야 합니다.

# --hints--

`funModule`는 정의되어야 하며 객체를 반환해야 합니다.

```js
assert(typeof funModule === 'object');
```

`funModule.isCuteMixin`는 함수를 접근해야 합니다.

```js
assert(typeof funModule.isCuteMixin === 'function');
```

`funModule.singMixin`는 함수를 접근해야 합니다.

```js
assert(typeof funModule.singMixin === 'function');
```

# --seed--

## --seed-contents--

```js
let isCuteMixin = function(obj) {
  obj.isCute = function() {
    return true;
  };
};
let singMixin = function(obj) {
  obj.sing = function() {
    console.log("Singing to an awesome tune");
  };
};
```

# --solutions--

```js
const funModule = (function () {
  return {
    isCuteMixin: obj => {
      obj.isCute = () => true;
    },
    singMixin: obj => {
      obj.sing = () => console.log("Singing to an awesome tune");
    }
  };
})();
```
