---
id: 587d7dae367417b2b2512b7b
title: 고유 프로퍼티 이해하기
challengeType: 1
forumTopicId: 301326
dashedName: understand-own-properties
---

# --description--

다음 예제에서 `Bird` 생성자는 두 개의 프로퍼티 `name`과 `numLegs`를 정의합니다:

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let duck = new Bird("Donald");
let canary = new Bird("Tweety");
```

`name`과 `numLegs`는 <dfn>고유 프로퍼티</dfn>라고 부르는데, 이는 인스턴스 객체에 직접 정의되었기 때문입니다. 즉, `duck`와 `canary` 각각은 이 프로퍼티들의 별도 복사본을 가지고 있습니다. 사실 `Bird`의 모든 인스턴스는 이 프로퍼티들의 고유 복사본을 갖게 됩니다. 다음 코드는 `duck`의 모든 고유 프로퍼티를 배열 `ownProps`에 추가합니다:

```js
let ownProps = [];

for (let property in duck) {
  if(duck.hasOwnProperty(property)) {
    ownProps.push(property);
  }
}

console.log(ownProps);
```

콘솔에는 값 `["name", "numLegs"]`이 출력될 것입니다.

# --instructions--

`canary`의 고유 프로퍼티를 배열 `ownProps`에 추가하세요.

# --hints--

`ownProps`에는 값 `numLegs`와 `name`가 포함되어야 합니다.

```js
assert(ownProps.indexOf('name') !== -1 && ownProps.indexOf('numLegs') !== -1);
```

내장 메서드 `Object.keys()`를 사용하지 않고 이 도전 과제를 해결하세요.

```js
assert(!/Object(\.keys|\[(['"`])keys\2\])/.test(__helpers.removeJSComments(code)));
```

`ownProps` 배열을 하드코딩하지 않고 이 도전 과제를 해결하세요.

```js
assert(
  !/\[\s*(?:'|")(?:name|numLegs)|(?:push|concat)\(\s*(?:'|")(?:name|numLegs)/.test(
    __helpers.removeJSComments(code)
  )
);
```

# --seed--

## --seed-contents--

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let canary = new Bird("Tweety");
let ownProps = [];
// Only change code below this line
```

# --solutions--

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let canary = new Bird("Tweety");
function getOwnProps (obj) {
  const props = [];

  for (let prop in obj) {
    if (obj.hasOwnProperty(prop)) {
      props.push(prop);
    }
  }

  return props;
}

const ownProps = getOwnProps(canary);
```
