---
id: 587d7dad367417b2b2512b77
title: 생성자 함수 정의하기
challengeType: 1
forumTopicId: 16804
dashedName: define-a-constructor-function
---

# --description--

<dfn>생성자</dfn>는 새로운 객체를 만드는 함수입니다. 생성자는 새 객체가 가질 속성과 동작을 정의합니다. 생성자는 새 객체를 만드는 청사진이라고 생각하세요.

다음은 생성자의 예입니다:

```js
function Bird() {
  this.name = "Albert";
  this.color = "blue";
  this.numLegs = 2;
}
```

이 생성자는 `Bird` 객체를 정의하며, 속성 `name`, `color`, `numLegs`는 각각 Albert, blue, 2로 설정됩니다. 생성자는 몇 가지 관례를 따릅니다:

<ul><li>생성자는 다른 함수와 구분하기 위해 이름을 대문자로 시작합니다.</li><li>생성자는 생성할 객체의 속성을 설정하기 위해 <code>this</code> 키워드를 사용합니다. 생성자 내부에서 <code>this</code>는 생성할 새 객체를 가리킵니다.</li><li>생성자는 다른 함수처럼 값을 반환하는 대신 속성과 동작을 정의합니다.</li></ul>

# --instructions--

속성 `Dog`, `name`, `color`가 각각 문자열, 문자열, 숫자로 설정된 생성자 `numLegs`를 만드세요.

# --hints--

`Dog`는 문자열로 설정된 `name` 속성을 가져야 합니다.

```js
assert(typeof new Dog().name === 'string');
```

`Dog`는 문자열로 설정된 `color` 속성을 가져야 합니다.

```js
assert(typeof new Dog().color === 'string');
```

`Dog`는 숫자로 설정된 `numLegs` 속성을 가져야 합니다.

```js
assert(typeof new Dog().numLegs === 'number');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function Dog (name, color, numLegs) {
  this.name = 'name';
  this.color = 'color';
  this.numLegs = 4;
}
```
