---
id: 587d7dae367417b2b2512b7c
title: 중복 코드를 줄이기 위해 프로토타입 속성 사용하기
challengeType: 1
forumTopicId: 301336
dashedName: use-prototype-properties-to-reduce-duplicate-code
---

# --description--

`numLegs`는 `Bird`의 모든 인스턴스에서 아마도 같은 값을 가지므로, 사실상 각 `numLegs` 인스턴스 안에 중복된 변수 `Bird`가 있는 셈입니다.

인스턴스가 두 개뿐이라면 문제가 되지 않을 수 있지만, 수백만 개의 인스턴스가 있다고 상상해 보세요. 그럼 중복된 변수가 엄청나게 많아집니다.

더 나은 방법은 `prototype`의 `Bird`를 사용하는 것입니다. `prototype`의 속성은 `Bird`의 모든 인스턴스가 공유합니다. `numLegs`에 `Bird prototype`를 추가하는 방법은 다음과 같습니다:

```js
Bird.prototype.numLegs = 2;
```

이제 `Bird`의 모든 인스턴스는 `numLegs` 속성을 갖게 됩니다.

```js
console.log(duck.numLegs);
console.log(canary.numLegs);
```

모든 인스턴스가 자동으로 `prototype`의 속성을 가지므로, `prototype`을 객체를 만드는 "레시피"라고 생각하세요. `prototype`와 `duck`의 `canary`는 `Bird` 생성자의 `Bird.prototype` 일부임을 주목하세요.

# --instructions--

`numLegs`의 `prototype` 프로토타입에 `Dog` 속성을 추가하세요.

# --hints--

`beagle`는 `numLegs` 속성을 가져야 합니다.

```js
assert(beagle.numLegs !== undefined);
```

`beagle.numLegs`는 숫자여야 합니다.

```js
assert(typeof beagle.numLegs === 'number');
```

`numLegs`는 자신의 속성이 아닌 `prototype` 속성이어야 합니다.

```js
assert(beagle.hasOwnProperty('numLegs') === false);
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}



// Only change code above this line
let beagle = new Dog("Snoopy");
```

# --solutions--

```js
function Dog (name) {
  this.name = name;
}
Dog.prototype.numLegs = 4;
let beagle = new Dog("Snoopy");
```
