---
id: 587d7dac367417b2b2512b73
title: 기본 JavaScript 객체 만들기
challengeType: 1
forumTopicId: 301317
dashedName: create-a-basic-javascript-object
---

# --description--

사람들이 매일 보는 것들, 예를 들어 자동차, 가게, 새를 생각해 보세요. 이들은 모두 <dfn>객체</dfn>입니다: 사람들이 관찰하고 상호작용할 수 있는 실체가 있는 것들입니다.

이 객체들의 어떤 특성들이 있을까요? 자동차는 바퀴가 있습니다. 가게는 물건을 팝니다. 새는 날개가 있습니다.

이러한 특성들, 즉 <dfn>속성</dfn>은 객체를 구성하는 요소를 정의합니다. 비슷한 객체들은 같은 속성을 공유하지만, 그 속성의 값은 다를 수 있다는 점에 주목하세요. 예를 들어, 모든 자동차는 바퀴가 있지만, 모든 자동차가 같은 수의 바퀴를 가진 것은 아닙니다.

JavaScript의 객체는 현실 세계의 객체를 모델링하는 데 사용되며, 현실 세계의 대응물처럼 속성과 동작을 부여합니다. 다음은 이러한 개념을 사용해 `duck` 객체를 만드는 예시입니다:

```js
let duck = {
  name: "Aflac",
  numLegs: 2
};
```

이 `duck` 객체는 `name`의 `Aflac`와 2의 `numLegs`라는 두 개의 속성/값 쌍을 가지고 있습니다.

# --instructions--

`dog` 객체를 만들고 `name`과 `numLegs` 속성을 각각 문자열과 숫자로 설정하세요.

# --hints--

`dog`는 객체여야 합니다.

```js
assert(typeof dog === 'object');
```

`dog`는 문자열로 설정된 `name` 속성을 가져야 합니다.

```js
assert(typeof dog.name === 'string');
```

`dog`는 숫자로 설정된 `numLegs` 속성을 가져야 합니다.

```js
assert(typeof dog.numLegs === 'number');
```

# --seed--

## --seed-contents--

```js
let dog = {

};
```

# --solutions--

```js
let dog = {
  name: '',
  numLegs: 4
};
```
