---
id: 587d7db0367417b2b2512b81
title: 객체의 프로토타입이 어디서 오는지 이해하기
challengeType: 1
forumTopicId: 301330
dashedName: understand-where-an-objects-prototype-comes-from
---

# --description--

사람이 부모로부터 유전자를 물려받는 것처럼, 객체도 그것을 만든 생성자 함수로부터 `prototype`를 직접 물려받습니다. 예를 들어, 여기서 `Bird` 생성자가 `duck` 객체를 만듭니다:

```js
function Bird(name) {
  this.name = name;
}

let duck = new Bird("Donald");
```

`duck`는 `prototype` 생성자 함수로부터 `Bird`를 물려받습니다. 이 관계는 `isPrototypeOf` 메서드로 확인할 수 있습니다:

```js
Bird.prototype.isPrototypeOf(duck);
```

이 코드는 `true`를 반환할 것입니다.

# --instructions--

`isPrototypeOf`를 사용해 `prototype`의 `beagle`를 확인하세요.

# --hints--

`Dog.prototype`가 `prototype`의 `beagle`임을 보여줘야 합니다.

```js
assert(/Dog\.prototype\.isPrototypeOf\(beagle\)/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

let beagle = new Dog("Snoopy");

// Only change code below this line
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
let beagle = new Dog("Snoopy");
Dog.prototype.isPrototypeOf(beagle);
```
