---
id: 587d7db0367417b2b2512b83
title: 상속을 사용해 중복을 피하세요
challengeType: 1
forumTopicId: 301334
dashedName: use-inheritance-so-you-dont-repeat-yourself
---

# --description--

프로그래밍에는 <dfn>중복 금지(DRY, Don't Repeat Yourself)</dfn>라는 원칙이 있습니다. 중복된 코드는 변경할 때 여러 곳을 고쳐야 하므로 문제가 됩니다. 이는 보통 프로그래머의 작업량을 늘리고 오류 발생 가능성을 높입니다.

아래 예시에서 `describe` 메서드는 `Bird`과 `Dog`가 공유하고 있음을 확인하세요:

```js
Bird.prototype = {
  constructor: Bird,
  describe: function() {
    console.log("My name is " + this.name);
  }
};

Dog.prototype = {
  constructor: Dog,
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

`describe` 메서드는 두 곳에서 반복되고 있습니다. `supertype`라는 `Animal` (또는 부모)를 만들어 DRY 원칙에 맞게 코드를 수정할 수 있습니다:

```js
function Animal() { };

Animal.prototype = {
  constructor: Animal, 
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

`Animal`이 `describe` 메서드를 포함하므로 `Bird`과 `Dog`에서는 제거할 수 있습니다:

```js
Bird.prototype = {
  constructor: Bird
};

Dog.prototype = {
  constructor: Dog
};
```

# --instructions--

`eat` 메서드는 `Cat`과 `Bear` 모두에서 반복되고 있습니다. `eat` 메서드를 `Animal` `supertype`로 옮겨 DRY 원칙에 맞게 코드를 수정하세요.

# --hints--

`Animal.prototype`은 `eat` 속성을 가져야 합니다.

```js
assert(Animal.prototype.hasOwnProperty('eat'));
```

`Bear.prototype`은 `eat` 속성을 가져서는 안 됩니다.

```js
assert(!Bear.prototype.hasOwnProperty('eat'));
```

`Cat.prototype`은 `eat` 속성을 가져서는 안 됩니다.

```js
assert(!Cat.prototype.hasOwnProperty('eat'));
```

# --seed--

## --seed-contents--

```js
function Cat(name) {
  this.name = name;
}

Cat.prototype = {
  constructor: Cat,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Bear(name) {
  this.name = name;
}

Bear.prototype = {
  constructor: Bear,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Animal() { }

Animal.prototype = {
  constructor: Animal,

};
```

# --solutions--

```js
function Cat(name) {
  this.name = name;
}

Cat.prototype = {
  constructor: Cat
};

function Bear(name) {
  this.name = name;
}

Bear.prototype = {
  constructor: Bear
};

function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};
```
