---
id: 587d7db2367417b2b2512b89
title: 관련 없는 객체들 사이에 공통 동작을 추가하기 위해 믹스인 사용하기
challengeType: 1
forumTopicId: 301331
dashedName: use-a-mixin-to-add-common-behavior-between-unrelated-objects
---

# --description--

지금까지 보았듯이 동작은 상속을 통해 공유됩니다. 하지만 상속이 최선의 해결책이 아닌 경우도 있습니다. 상속은 `Bird`와 `Airplane`처럼 관련 없는 객체에는 잘 맞지 않습니다. 두 객체 모두 날 수 있지만 `Bird`는 `Airplane`의 유형이 아니며 그 반대도 마찬가지입니다.

관련 없는 객체에는 <dfn>믹스인</dfn>을 사용하는 것이 더 좋습니다. 믹스인은 다른 객체가 함수 모음을 사용할 수 있게 합니다.

```js
let flyMixin = function(obj) {
  obj.fly = function() {
    console.log("Flying, wooosh!");
  }
};
```

`flyMixin`는 어떤 객체든 받아서 `fly` 메서드를 부여합니다.

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let plane = {
  model: "777",
  numPassengers: 524
};

flyMixin(bird);
flyMixin(plane);
```

여기서 `bird`와 `plane`가 `flyMixin`에 전달되고, `flyMixin`는 `fly` 함수를 각 객체에 할당합니다. 이제 `bird`와 `plane` 모두 날 수 있습니다:

```js
bird.fly();
plane.fly();
```

콘솔에는 `Flying, wooosh!` 문자열이 두 번 출력되며, 각각 `.fly()` 호출에 해당합니다.

믹스인이 관련 없는 객체 `bird`와 `plane`가 같은 `fly` 메서드를 재사용할 수 있게 하는 점에 주목하세요.

# --instructions--

`glideMixin`라는 이름의 믹스인을 만들어 `glide`라는 메서드를 정의하세요. 그런 다음 `glideMixin`를 사용해 `bird`와 `boat` 모두에게 활공하는 능력을 부여하세요.

# --hints--

코드는 함수인 `glideMixin` 변수를 선언해야 합니다.

```js
assert(typeof glideMixin === 'function');
```

코드는 `bird` 객체에 `glide` 메서드를 부여하기 위해 `glideMixin`를 사용해야 합니다.

```js
assert(typeof bird.glide === 'function');
```

코드는 `boat` 객체에 `glide` 메서드를 부여하기 위해 `glideMixin`를 사용해야 합니다.

```js
assert(typeof boat.glide === 'function');
```

# --seed--

## --seed-contents--

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let boat = {
  name: "Warrior",
  type: "race-boat"
};

// Only change code below this line
```

# --solutions--

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let boat = {
  name: "Warrior",
  type: "race-boat"
};
function glideMixin (obj) {
  obj.glide = () => 'Gliding!';
}

glideMixin(bird);
glideMixin(boat);
```
