---
id: 587d7db1367417b2b2512b85
title: 자식 프로토타입을 부모 인스턴스로 설정하기
challengeType: 1
forumTopicId: 301325
dashedName: set-the-childs-prototype-to-an-instance-of-the-parent
---

# --description--

이전 도전 과제에서는 상위 타입(또는 부모) `Animal`에서 동작을 상속받는 첫 번째 단계인 `Animal`의 새 인스턴스를 만드는 방법을 배웠습니다.

이번 도전 과제에서는 다음 단계인 하위 타입(또는 자식) `prototype`의 `Bird`를 `Animal`의 인스턴스로 설정하는 방법을 다룹니다.

```js
Bird.prototype = Object.create(Animal.prototype);
```

`prototype`는 객체를 만드는 "레시피"와 같다는 점을 기억하세요. 어떤 면에서는 `Bird`의 레시피가 이제 `Animal`의 모든 핵심 "재료"를 포함하게 됩니다.

```js
let duck = new Bird("Donald");
duck.eat();
```

`duck`는 `Animal`의 모든 프로퍼티를 상속받으며, 여기에는 `eat` 메서드도 포함됩니다.

# --instructions--

`Dog`의 인스턴스가 `Animal`로부터 상속받도록 코드를 수정하세요.

# --hints--

`Dog.prototype`는 `Animal`의 인스턴스여야 합니다.

```js
assert(Animal.prototype.isPrototypeOf(Dog.prototype));
```

# --seed--

## --seed-contents--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Dog() { }

// Only change code below this line


let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Dog() { }
Dog.prototype = Object.create(Animal.prototype);

let beagle = new Dog();
beagle.eat();
```
