---
id: 66bb3e20d3dc5b6d0a21f5dd
title: Echtzeit-Zähler erstellen
challengeType: 25
dashedName: build-a-real-time-counter
demoType: onClick
---

# --description--

In diesem Labor erstellen Sie einen Echtzeit-Zeichenzähler mit JavaScript. Der Zeichenzähler zeigt die Anzahl der eingegebenen Zeichen in einem `textarea`-Element an. Der Zähler wird in Echtzeit aktualisiert, während der Benutzer im `textarea` tippt.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten ein `textarea`-Element mit dem `id` von `text-input` haben.
2. Es sollte ein `p`-Element mit dem `id` von `char-count` geben. Es sollte anfangs den Text `Character Count: 0/50` enthalten. Dieser Platzhaltertext soll durch die aktuelle Anzahl der Zeichen im `textarea`-Element ersetzt werden.
3. Wenn das `#text-input`-Element den Text `hello coder` enthält, soll das `#char-count`-Element den Text `"Character Count: 11/50"` enthalten.
4. Wenn das `#text-input`-Element den Text `i am a programmer` enthält, soll das `#char-count`-Element den Text `"Character Count: 17/50"` enthalten.
5. Wenn das `#text-input`-Element den Text `hello world` enthält, soll das `#char-count`-Element den Text `"Character Count: 11/50"` enthalten.
6. Wenn das `#text-input`-Element den Text `I am learning a new language and it is called c++.` enthält, soll das `#char-count`-Element den Text `"Character Count: 50/50"` enthalten. **HINWEIS**: Obwohl das `maxlength`-Attribut die korrekte Funktionalität erreichen würde, erfordert dieses Labor eine rein JavaScript-basierte Lösung.  
7. Der Benutzer sollte nicht mehr als `50` Zeichen im `textarea`-Element eingeben können. Wenn die Zeichenzahl `50` erreicht, soll jede zusätzliche Eingabe automatisch abgeschnitten werden und der Text `Character Count: 50/50` im `red` angezeigt werden.

**Hinweis:** Vergewissern Sie sich, dass Sie Ihr Stylesheet und die JavaScript-Datei in Ihrem HTML verlinkt haben.

# --hints--

Sie sollten ein `textarea`-Element mit dem `id` von `text-input` haben.

```js
const textInput = document.querySelector('textarea');
assert.equal(textInput.id, 'text-input');
```

Sie sollten ein `p`-Element mit dem `id` von `char-count` haben.

```js
assert.exists("p#char-count");
```

Das `#char-count`-Element sollte anfangs den Text `Character Count: 0/50` enthalten.

```js
const charCount = document.querySelector('p#char-count')?.innerText;
assert.equal(charCount, 'Character Count: 0/50');
```

Wenn das `#text-input`-Element den Text `hello coder` enthält, sollte das `#char-count`-Element den Text `"Character Count: 11/50"` enthalten.

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "hello coder";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(charCount.innerText.trim(), 'Character Count: 11/50');
```

Wenn die Zeichenzahl `50` ist, sollte der Text im `red` angezeigt werden.

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "I am learning a new language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
const computedStyle = window.getComputedStyle(charCount);
assert.oneOf(computedStyle.color, [
  "red",
  "rgb(255, 0, 0)",
  "rgb(100%, 0%, 0%)",
  "rgba(255, 0, 0, 1)",
  "rgba(100%, 0%, 0%, 1)",
  "hsl(0, 100%, 50%)",
  "hsla(0, 100%, 50%, 1)"
]);
```

Wenn die Zeichenzahl größer oder gleich `50` ist, sollte die zusätzliche Eingabe abgeschnitten werden.

```js
const textInput = document.getElementById('text-input');
textInput.value = "I am learning a new programming language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(textInput.value.length, 50);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real Time Counter</title>
</head>

<body>

</body>


</html>
```

```css

```

```js
    
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real-time Character Counter</title>
    <link rel="stylesheet" href="styles.css">

<body>
    <div class="container">
        <h1>Real-Time Character Counter</h1>
        <textarea id="text-input" placeholder="Type something..."></textarea>
        <p id="char-count">Character Count: 0/50</p>
    </div>

<script src="script.js"></script>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    font-family: Arial, sans-serif;
    background-color: #c0c1ec;
    margin: 0;
}

h1{
    background-color: #f2f48b;
    padding-left: 10px;
    padding-right: 10px;
}
.container {
    text-align: center;
}

textarea {
    width: 300px;
    height: 100px;
    font-size: 18px;
    padding: 10px;
    border: 1px solid #ccc;
    border-radius: 5px;
    resize: none;
    display: block;
    margin: 0 auto;
    border: 1px solid slategray;
    box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

textarea:focus {
    outline: none;
    border-color: #f2f48b;
}

#char-count {
    margin-top: 10px;
    font-size: 18px;
    color: darkslategray;
    font-weight: bold;
}
```

```js
document.getElementById("text-input").addEventListener("input", function () {
  let textInput = document.getElementById("text-input");
  let charCount = textInput.value.length;

  if (charCount > 50) {
    textInput.value = textInput.value.substring(0, 50);
    charCount = 50;
  }

  let charCountDisplay = document.getElementById("char-count");
  charCountDisplay.innerHTML = `Character Count: ${charCount}/50`;

  if (charCount === 50) {
    charCountDisplay.style.color = "red";
  } 

});
```
