---
id: 5a24c314108439a4d4036149
title: е°ҮеҚҖеҹҹзҡ„зӢҖж…ӢжҠҪеҸ–еҲ° Redux
challengeType: 6
forumTopicId: 301428
dashedName: extract-local-state-into-redux
---

# --description--

дҪ еҝ«е®ҢжҲҗдәҶпјҒи«ӢиЁҳеҫ—дҪ ж’°еҜ«дәҶжүҖжңү Redux зЁӢејҸзўјпјҢи®“ Redux иғҪеӨ жҺ§еҲ¶дҪ зҡ„ React иЁҠжҒҜжҮүз”ЁзЁӢејҸзҡ„зӢҖж…Ӣз®ЎзҗҶгҖӮж—ўз„¶ Redux е·Із¶“йҖЈжҺҘпјҢдҪ йңҖиҰҒе°ҮзӢҖж…Ӣз®ЎзҗҶеҫһ `Presentational` зө„д»¶дёӯжҠҪйӣўдёҰз§»еҲ° ReduxгҖӮзҸҫеңЁдҪ е·Із¶“йҖЈжҺҘ ReduxпјҢдҪҶдҪ д»ҚеңЁ `Presentational` зө„д»¶е…§йғЁеҚҖеҹҹең°иҷ•зҗҶзӢҖж…ӢгҖӮ

# --instructions--

еңЁ `Presentational` зө„д»¶дёӯпјҢйҰ–е…ҲпјҢз§»йҷӨеҚҖеҹҹзҡ„ `state` дёӯзҡ„ `messages` еұ¬жҖ§гҖӮйҖҷдәӣиЁҠжҒҜе°Үз”ұ Redux з®ЎзҗҶгҖӮжҺҘи‘—пјҢдҝ®ж”№ `submitMessage()` ж–№жі•пјҢдҪҝе…¶еҲҶжҙҫдҫҶиҮӘ `this.props` зҡ„ `submitNewMessage()`пјҢдёҰе°ҮеҚҖеҹҹзҡ„ `state` дёӯзӣ®еүҚзҡ„иЁҠжҒҜијёе…ҘдҪңзӮәеј•ж•ёеӮіе…ҘгҖӮеӣ зӮәдҪ е·ІеҫһеҚҖеҹҹзҡ„ state дёӯз§»йҷӨ `messages`пјҢжүҖд»Ҙд№ҹеҫһйҖҷиЈЎе°Қ `this.setState()` зҡ„е‘јеҸ«дёӯз§»йҷӨ `messages` еұ¬жҖ§гҖӮжңҖеҫҢпјҢдҝ®ж”№ `render()` ж–№жі•пјҢдҪҝе…¶е°Қеҫһ `props` жҺҘж”¶еҲ°зҡ„иЁҠжҒҜйҖІиЎҢжҳ е°„пјҢиҖҢйқһеҫһ `state`гҖӮ

дёҖж—Ұе®ҢжҲҗйҖҷдәӣжӣҙж”№пјҢжҮүз”ЁзЁӢејҸе°Үз№јзәҢд»ҘзӣёеҗҢж–№ејҸйҒӢдҪңпјҢйҷӨдәҶз”ұ Redux з®ЎзҗҶзӢҖж…ӢгҖӮйҖҷеҖӢзҜ„дҫӢд№ҹиӘӘжҳҺдәҶзө„д»¶еҸҜиғҪж“ҒжңүеҚҖеҹҹзҡ„ `state`пјҡдҪ зҡ„зө„д»¶д»Қз„¶еңЁиҮӘе·ұзҡ„ `state` дёӯеҚҖеҹҹжҖ§ең°иҝҪи№ӨдҪҝз”ЁиҖ…ијёе…ҘгҖӮдҪ еҸҜд»ҘзңӢеҲ° Redux еҰӮдҪ•еңЁ React д№ӢдёҠжҸҗдҫӣдёҖеҖӢжңүз”Ёзҡ„зӢҖж…Ӣз®ЎзҗҶжЎҶжһ¶гҖӮдҪ жңҖеҲқеҸӘдҪҝз”Ё React зҡ„еҚҖеҹҹ state е°ұйҒ”жҲҗдәҶзӣёеҗҢзҡ„зөҗжһңпјҢйҖҷйҖҡеёёйҒ©з”Ёж–јз°Ўе–®зҡ„жҮүз”ЁзЁӢејҸгҖӮз„¶иҖҢпјҢйҡЁи‘—дҪ зҡ„жҮүз”ЁзЁӢејҸи®Ҡеҫ—жӣҙеӨ§дё”жӣҙиӨҮйӣңпјҢзӢҖж…Ӣз®ЎзҗҶд№ҹжңғи®Ҡеҫ—жӣҙиӨҮйӣңпјҢиҖҢйҖҷжӯЈжҳҜ Redux и§Јжұәзҡ„е•ҸйЎҢгҖӮ

# --hints--

`AppWrapper` жҮүи©Іе‘ҲзҸҫеҲ°й ҒйқўдёҠгҖӮ

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return mockedComponent.find('AppWrapper').length === 1;
  })()
);
```

`Presentational` зө„д»¶жҮүи©Іе‘ҲзҸҫеҲ°й ҒйқўдёҠгҖӮ

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return mockedComponent.find('Presentational').length === 1;
  })()
);
```

`Presentational` зө„д»¶жҮүи©Іе‘ҲзҸҫ `h2`гҖҒ`input`гҖҒ`button` е’Ң `ul` е…ғзҙ гҖӮ

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const PresentationalComponent = mockedComponent.find('Presentational');
    return (
      PresentationalComponent.find('div').length === 1 &&
      PresentationalComponent.find('h2').length === 1 &&
      PresentationalComponent.find('button').length === 1 &&
      PresentationalComponent.find('ul').length === 1
    );
  })()
);
```

`Presentational` зө„д»¶жҮүи©Іеҫһ Redux е„Іеӯҳеә«жҺҘж”¶ `messages` дҪңзӮәеұ¬жҖ§гҖӮ

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const PresentationalComponent = mockedComponent.find('Presentational');
    const props = PresentationalComponent.props();
    return Array.isArray(props.messages);
  })()
);
```

`Presentational` зө„д»¶жҮүи©ІжҺҘж”¶ `submitMessage` action creator дҪңзӮәеұ¬жҖ§гҖӮ

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const PresentationalComponent = mockedComponent.find('Presentational');
    const props = PresentationalComponent.props();
    return typeof props.submitNewMessage === 'function';
  })()
);
```

`Presentational` зө„д»¶зҡ„зӢҖж…ӢжҮүи©ІеҢ…еҗ«дёҖеҖӢеұ¬жҖ§ `input`пјҢи©Іеұ¬жҖ§еҲқе§ӢеҢ–зӮәз©әеӯ—дёІгҖӮ

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const PresentationalState = mockedComponent
      .find('Presentational')
      .instance().state;
    return (
      typeof PresentationalState.input === 'string' &&
      Object.keys(PresentationalState).length === 1
    );
  })()
);
```

еңЁ `input` е…ғзҙ дёӯијёе…ҘжҮүи©Іжңғжӣҙж–° `Presentational` зө„д»¶зҡ„зӢҖж…ӢгҖӮ

```js
  const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
  const testValue = '__MOCK__INPUT__';
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 100));
  const causeChange = (c, v) =>
    c.find('input').simulate('change', { target: { value: v } });
  let initialInput = mockedComponent.find('Presentational').find('input');
  const changed = () => {
    causeChange(mockedComponent, testValue);
    return waitForIt(() => mockedComponent);
  };
  const updated = await changed();
  const updatedInput = updated.find('Presentational').find('input');
  assert(
    initialInput.props().value === '' &&
      updatedInput.props().value === '__MOCK__INPUT__'
  );
```

еңЁ `Presentational` зө„д»¶дёҠеҲҶжҙҫ `submitMessage` жҮүи©Іжңғжӣҙж–° Redux е„ІеӯҳдёҰжё…йҷӨеҚҖеҹҹзҡ„зӢҖж…Ӣдёӯзҡ„ијёе…ҘгҖӮ

```js
  const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 100));
  let beforeProps = mockedComponent.find('Presentational').props();
  const testValue = '__TEST__EVENT__INPUT__';
  const causeChange = (c, v) =>
    c.find('input').simulate('change', { target: { value: v } });
  const changed = () => {
    causeChange(mockedComponent, testValue);
    return waitForIt(() => mockedComponent);
  };
  const clickButton = () => {
    mockedComponent.find('button').simulate('click');
    return waitForIt(() => mockedComponent);
  };
  const afterChange = await changed();
  const afterChangeInput = afterChange.find('input').props().value;
  const afterClick = await clickButton();
  const afterProps = mockedComponent.find('Presentational').props();
  assert(
    beforeProps.messages.length === 0 &&
      afterChangeInput === testValue &&
      afterProps.messages.pop() === testValue &&
      afterClick.find('input').props().value === ''
  );
```

`Presentational` зө„д»¶жҮүи©Іе‘ҲзҸҫдҫҶиҮӘ Redux е„Іеӯҳеә«зҡ„ `messages`гҖӮ

```js
  const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 100));
  let beforeProps = mockedComponent.find('Presentational').props();
  const testValue = '__TEST__EVENT__INPUT__';
  const causeChange = (c, v) =>
    c.find('input').simulate('change', { target: { value: v } });
  const changed = () => {
    causeChange(mockedComponent, testValue);
    return waitForIt(() => mockedComponent);
  };
  const clickButton = () => {
    mockedComponent.find('button').simulate('click');
    return waitForIt(() => mockedComponent);
  };
  const afterChange = await changed();
  const afterChangeInput = afterChange.find('input').props().value;
  const afterClick = await clickButton();
  const afterProps = mockedComponent.find('Presentational').props();
  assert(
    beforeProps.messages.length === 0 &&
      afterChangeInput === testValue &&
      afterProps.messages.pop() === testValue &&
      afterClick.find('input').props().value === '' &&
      afterClick.find('ul').childAt(0).text() === testValue
  );
```

# --seed--

## --seed-contents--

```jsx
// Redux:
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message: message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};

const store = Redux.createStore(messageReducer);

// React:
const Provider = ReactRedux.Provider;
const connect = ReactRedux.connect;

// Change code below this line
class Presentational extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
    this.handleChange = this.handleChange.bind(this);
    this.submitMessage = this.submitMessage.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  submitMessage() {
    this.setState((state) => ({
      input: '',
      messages: state.messages.concat(state.input)
    }));
  }
  render() {
    return (
      <div>
        <h2>Type in a new Message:</h2>
        <input
          value={this.state.input}
          onChange={this.handleChange}/><br/>
        <button onClick={this.submitMessage}>Submit</button>
        <ul>
          {this.state.messages.map( (message, idx) => {
              return (
                 <li key={idx}>{message}</li>
              )
            })
          }
        </ul>
      </div>
    );
  }
};
// Change code above this line

const mapStateToProps = (state) => {
  return {messages: state}
};

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: (message) => {
      dispatch(addMessage(message))
    }
  }
};

const Container = connect(mapStateToProps, mapDispatchToProps)(Presentational);

class AppWrapper extends React.Component {
  render() {
    return (
      <Provider store={store}>
        <Container/>
      </Provider>
    );
  }
};
```

# --solutions--

```jsx
// Redux:
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message: message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};

const store = Redux.createStore(messageReducer);

// React:
const Provider = ReactRedux.Provider;
const connect = ReactRedux.connect;

// Change code below this line
class Presentational extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    }
 this.handleChange = this.handleChange.bind(this);
 this.submitMessage = this.submitMessage.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  submitMessage() {
    this.props.submitNewMessage(this.state.input);
    this.setState({
      input: ''
    });
  }
  render() {
    return (
      <div>
        <h2>Type in a new Message:</h2>
        <input
          value={this.state.input}
          onChange={this.handleChange}/><br/>
        <button onClick={this.submitMessage}>Submit</button>
        <ul>
          {this.props.messages.map( (message, idx) => {
              return (
                 <li key={idx}>{message}</li>
              )
            })
          }
        </ul>
      </div>
    );
  }
};
// Change code above this line

const mapStateToProps = (state) => {
  return {messages: state}
};

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: (message) => {
      dispatch(addMessage(message))
    }
  }
};

const Container = connect(mapStateToProps, mapDispatchToProps)(Presentational);

class AppWrapper extends React.Component {
  render() {
    return (
      <Provider store={store}>
        <Container/>
      </Provider>
    );
  }
};
```
