---
id: 5a24c314108439a4d403614a
title: 從此繼續前進
challengeType: 6
forumTopicId: 301434
dashedName: moving-forward-from-here
---

# --description--

恭喜你！你完成了有關 React 和 Redux 的課程。在你繼續之前，有最後一個值得指出的條款。通常，你不會在像這樣的程式碼編輯器中撰寫 React 應用程式。這個挑戰讓你一窺如果你在自己的機器上使用檔案系統時語法的樣貌。程式碼應該看起來相似，除了使用 `import` 述語（這些會匯入在挑戰中為你提供的所有相依性）。

最後，撰寫 React 和 Redux 程式碼通常需要一些設定。這可能很快變得複雜。如果你有興趣在自己的機器上進行實驗，<a href="https://www.freecodecamp.org/news/install-react-with-create-react-app/" target="_blank" rel="noopener noreferrer nofollow">Create React App</a> 已經組態完成並可立即使用。

或者，你可以在 CodePen 中啟用 Babel 作為 JavaScript 前處理器，將 React 和 ReactDOM 添加為外部 JavaScript 資源，並且也可以在那裡進行開發。

# --instructions--

將訊息 `'Now I know React and Redux!'` 記錄到主控台。

# --hints--

`Now I know React and Redux!` 訊息應該被記錄到主控台。

```js
assert(
    /console\s*\.\s*log\s*\(\s*('|"|`)Now I know React and Redux!\1\s*\)/.test(
      code)
  );
```

# --seed--

## --seed-contents--

```jsx
/*
import React from 'react'
import ReactDOM from 'react-dom'
import { Provider, connect } from 'react-redux'
import { createStore, combineReducers, applyMiddleware } from 'redux'
import thunk from 'redux-thunk'

import rootReducer from './redux/reducers'
import App from './components/App'

const store = createStore(
  rootReducer,
  applyMiddleware(thunk)
);

ReactDOM.render(
  <Provider store={store}>
    <App/>
  </Provider>,
  document.getElementById('root')
);
*/

// Only change code below this line
```

# --solutions--

```jsx
console.log('Now I know React and Redux!');
```
