---
id: 5a24c314108439a4d4036145
title: 將 State 映射到 Props
challengeType: 6
forumTopicId: 301433
dashedName: map-state-to-props
---

# --description--

`Provider` 組件允許你為你的 React 組件提供 `state` 和 `dispatch`，但你必須明確指定你想要的狀態和動作。這樣，你就能確保每個組件只存取它所需要的狀態。你可以透過建立兩個函式（程式）：`mapStateToProps()` 和 `mapDispatchToProps()` 來達成這個目的。

在這些函式中，你宣告你想要存取的狀態部分以及你需要能夠分派的動作建立者。當這些函式就緒後，你將會看到如何使用 React Redux 的 `connect` 方法，將它們連接到你的組件中，在另一個挑戰中。

**注意：** 在背後，React Redux 使用 `store.subscribe()` 方法來實作 `mapStateToProps()`。

# --instructions--

創建一個函式（程式）`mapStateToProps()`。這個函式（程式）應該接受 `state` 作為引數，然後傳回一個物件，該物件將該狀態映射到特定的屬性名稱。這些屬性將透過 `props` 變得可供你的組件存取。由於此範例將應用程式的整個狀態保存在單一陣列中，你可以將該整個狀態傳遞給你的組件。在被傳回的物件中創建一個屬性 `messages`，並將其設定為 `state`。

# --hints--

const `state` 應該是一個空的陣列。

```js
assert(Array.isArray(state) && state.length === 0);
```

`mapStateToProps` 應該是一個函式（程式）。

```js
assert(typeof mapStateToProps === 'function');
```

`mapStateToProps` 應該傳回一個物件。

```js
assert(typeof mapStateToProps() === 'object');
```

將陣列作為狀態傳遞給 `mapStateToProps`，應該傳回此陣列並指定給 `messages` 的鍵。

```js
assert(mapStateToProps(['messages']).messages.pop() === 'messages');
```

# --seed--

## --seed-contents--

```jsx
const state = [];

// Change code below this line
```

# --solutions--

```jsx
const state = [];

// Change code below this line

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};
```
