---
id: 5a24c314108439a4d4036144
title: 使用 Provider 將 Redux 連接到 React
challengeType: 6
forumTopicId: 301435
dashedName: use-provider-to-connect-redux-to-react
---

# --description--

在上一個挑戰中，你創建了一個 Redux store 來控制代碼 `messages` 陣列，並創建了一個用於添加新訊息的 action。下一步是為 React 提供存取 Redux store 及其需要分派更新的 actions。React Redux 提供了它的 `react-redux` 軟體包來協助完成這些任務。

React Redux 提供一個包含兩個關鍵特性的簡易 API：`Provider` 和 `connect`。另一個挑戰涵蓋了 `connect`。`Provider` 是 React Redux 的一個外覆組件，它包裝你的 React 應用程式。這個外覆組件讓你能在整個組件樹中存取 Redux 的 `store` 和 `dispatch` 函式。`Provider` 接受兩個屬性，Redux store 和你的應用程式的子組件。為 App 組件定義 `Provider` 可能看起來像這樣：

```jsx
<Provider store={store}>
  <App/>
</Provider>
```

# --instructions--

程式碼編輯器現在顯示你過去幾個挑戰中的所有 Redux 和 React 程式碼。它包含 Redux store、actions 和 `DisplayMessages` 組件。唯一新增的部分是底部的 `AppWrapper` 組件。使用這個頂層組件來算繪來自 `ReactRedux` 的 `Provider`，並將 Redux store 作為屬性傳遞。然後將 `DisplayMessages` 組件作為子組件算繪。完成後，你應該會看到你的 React 組件呈現在頁面上。

**注意：** React Redux 在此處以全域變數形式提供，因此你可以使用點記法存取 `Provider`。編輯器中的程式碼利用了這點，並將其設定為常數 `Provider`，供你在 `AppWrapper` 呈現方法中使用。

# --hints--

`AppWrapper` 應該要呈現。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return mockedComponent.find('AppWrapper').length === 1;
  })()
);
```

`Provider` 外覆組件應該傳入一個 `store` 屬性，該屬性與 Redux store 相等。

```js
assert(
    (function () {
      const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
      return __helpers
        .removeWhiteSpace(code)
        .includes('<Providerstore={store}>');
    })()
  );
```

`DisplayMessages` 應該作為 `AppWrapper` 的子元件來呈現。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return (
      mockedComponent.find('AppWrapper').find('DisplayMessages').length === 1
    );
  })()
);
```

`DisplayMessages` 組件應該呈現一個 `h2`、`input`、`button` 和 `ul` 元素。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h2').length === 1 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('ul').length === 1
    );
  })()
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<AppWrapper />, document.getElementById('root'))
```

## --seed-contents--

```jsx
// Redux:
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};



const store = Redux.createStore(messageReducer);

// React:

class DisplayMessages extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
    this.handleChange = this.handleChange.bind(this);
    this.submitMessage = this.submitMessage.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  submitMessage() {  
    this.setState((state) => {
      const currentMessage = state.input;
      return {
        input: '',
        messages: state.messages.concat(currentMessage)
      };
    });
  }
  render() {
    return (
      <div>
        <h2>Type in a new Message:</h2>
        <input
          value={this.state.input}
          onChange={this.handleChange}/><br/>
        <button onClick={this.submitMessage}>Submit</button>
        <ul>
          {this.state.messages.map( (message, idx) => {
              return (
                 <li key={idx}>{message}</li>
              )
            })
          }
        </ul>
      </div>
    );
  }
};

const Provider = ReactRedux.Provider;

class AppWrapper extends React.Component {
  // Render the Provider below this line

  // Change code above this line
};
```

# --solutions--

```jsx
// Redux:
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};

const store = Redux.createStore(messageReducer);

// React:

class DisplayMessages extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
 this.handleChange = this.handleChange.bind(this);
 this.submitMessage = this.submitMessage.bind(this);
 }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  submitMessage() {
    this.setState((state) => {
      const currentMessage = state.input;
      return {
        input: '',
        messages: state.messages.concat(currentMessage)
      };  
    });
  }
  render() {
    return (
      <div>
        <h2>Type in a new Message:</h2>
        <input
          value={this.state.input}
          onChange={this.handleChange}/><br/>
        <button onClick={this.submitMessage}>Submit</button>
        <ul>
          {this.state.messages.map( (message, idx) => {
              return (
                 <li key={idx}>{message}</li>
              )
            })
          }
        </ul>
      </div>
    );
  }
};

const Provider = ReactRedux.Provider;

class AppWrapper extends React.Component {
  // Change code below this line
  render() {
    return (
      <Provider store = {store}>
        <DisplayMessages/>
      </Provider>
    );
  }
  // Change code above this line
};
```
