---
id: 5a24c314108439a4d4036147
title: 將 Redux 連接到 React
challengeType: 6
forumTopicId: 301426
dashedName: connect-redux-to-react
---

# --description--

現在你已經撰寫了 `mapStateToProps()` 和 `mapDispatchToProps()` 兩個函式（程式），你可以使用它們將 `state` 和 `dispatch` 映射到你其中一個 React 組件的 `props`。React Redux 的 `connect` 方法可以處理這項任務。此方法接受兩個選用引數，分別是 `mapStateToProps()` 和 `mapDispatchToProps()`。它們是選用的，因為你可能有一個組件只需要存取 `state`，但不需要分派任何動作，或反之亦然。

要使用此方法，將函式作為引數傳入，並立即用你的組件呼叫結果。這種語法有點不尋常，看起來像：

```js
connect(mapStateToProps, mapDispatchToProps)(MyComponent)
```

**注意：** 如果你想省略 `connect` 方法的一個引數，你可以傳入 `null` 作為替代。

# --instructions--

程式碼編輯器有 `mapStateToProps()` 和 `mapDispatchToProps()` 函式，以及一個新的 React 組件叫做 `Presentational`。使用來自 `ReactRedux` 全域物件的 `connect` 方法將此組件連接到 Redux，並立即在 `Presentational` 組件上呼叫它。將結果指定給一個新的 `const`，名為 `ConnectedComponent`，用來表述已連接的組件。就是這樣，現在你已經連接到 Redux！嘗試將 `connect` 的任一引數改為 `null`，並觀察測試結果。

# --hints--

`Presentational` 組件應該要呈現。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return mockedComponent.find('Presentational').length === 1;
  })()
);
```

`Presentational` 組件應該透過 `connect` 接收一個屬性 `messages`。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const props = mockedComponent.find('Presentational').props();
    return props.messages === '__INITIAL__STATE__';
  })()
);
```

`Presentational` 組件應該透過 `connect` 接收一個 prop `submitNewMessage`。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const props = mockedComponent.find('Presentational').props();
    return typeof props.submitNewMessage === 'function';
  })()
);
```

# --seed--

## --after-user-code--

```jsx
const store = Redux.createStore(
  (state = '__INITIAL__STATE__', action) => state
);
class AppWrapper extends React.Component {
  render() {
    return (
      <ReactRedux.Provider store = {store}>
        <ConnectedComponent/>
      </ReactRedux.Provider>
    );
  }
};
ReactDOM.render(<AppWrapper />, document.getElementById('root'))
```

## --seed-contents--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: (message) => {
      dispatch(addMessage(message));
    }
  }
};

class Presentational extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h3>This is a Presentational Component</h3>
  }
};

const connect = ReactRedux.connect;
// Change code below this line
```

# --solutions--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: (message) => {
      dispatch(addMessage(message));
    }
  }
};

class Presentational extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h3>This is a Presentational Component</h3>
  }
};

const connect = ReactRedux.connect;
// Change code below this line

const ConnectedComponent = connect(mapStateToProps, mapDispatchToProps)(Presentational);
```
