---
id: 5a24c314108439a4d4036143
title: 將狀態邏輯抽取到 Redux
challengeType: 6
forumTopicId: 301429
dashedName: extract-state-logic-to-redux
---

# --description--

現在你完成了 React 組件，你需要將它在 `state` 中本地執行的邏輯移動到 Redux。這是將簡單 React 應用程式連接到 Redux 的第一步。你的應用程式唯一的功能是將使用者的新訊息添加到無序列表中。這個範例很簡單，以示範 React 和 Redux 如何協同運作。

# --instructions--

首先，定義一個動作型別 `ADD` 並將其設定為常數 `ADD`。接著，定義一個動作產生器 `addMessage()`，用來產生新增訊息的動作。你需要傳入一個 `message` 給這個動作產生器，並將該訊息包含在傳回的 `action` 中。

然後創建一個名為 `messageReducer()` 的 reducer，用來控制訊息的狀態。初始狀態應該相等於一個空的陣列。這個 reducer 應該將訊息添加到狀態中持有的訊息陣列，或傳回目前的狀態。最後，創建你的 Redux store 並將 reducer 傳入。

# --hints--

const `ADD` 應該存在並且持有一個值，該值與字串 `ADD` 相等。

```js
assert(ADD === 'ADD');
```

動作建立者 `addMessage` 應該傳回一個 `type` 相等於 `ADD` 且 `message` 相等於所傳入訊息的物件。

```js
assert(
  (function () {
    const addAction = addMessage('__TEST__MESSAGE__');
    return addAction.type === ADD && addAction.message === '__TEST__MESSAGE__';
  })()
);
```

`messageReducer` 應該是一個函式（程式）。

```js
assert(typeof messageReducer === 'function');
```

該 store 應該存在並且初始狀態設定為空陣列。

```js
assert(
  (function () {
    const initialState = store.getState();
    return typeof store === 'object' && initialState.length === 0;
  })()
);
```

對 store 分派 `addMessage` 應該會以不變更的方式將新的訊息添加到狀態中所持有的訊息陣列。

```js
assert(
  (function () {
    const initialState = store.getState();
    const isFrozen = DeepFreeze(initialState);
    store.dispatch(addMessage('__A__TEST__MESSAGE'));
    const addState = store.getState();
    return isFrozen && addState[0] === '__A__TEST__MESSAGE';
  })()
);
```

如果使用任何其他動作呼叫，`messageReducer` 應該傳回目前的狀態。

```js
assert(
  (function () {
    const addState = store.getState();
    store.dispatch({ type: 'FAKE_ACTION' });
    const testState = store.getState();
    return addState === testState;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Define ADD, addMessage(), messageReducer(), and store here:
```

# --solutions--

```jsx
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};

const store = Redux.createStore(messageReducer);
```
