---
id: bad87fee1348bd9aec908857
title: Tumia maoni kufafanua msimbo
challengeType: 0
forumTopicId: 18347
dashedName: use-comments-to-clarify-code
---

# --description--

Tunapoanza kutumia jQuery, tutabadilisha vipengele vya HTML bila haja ya kubadilisha moja kwa moja katika HTML.

Hakikisha kila mtu anajua kwamba hawapaswi kubadilisha msimbo huu moja kwa moja.

Kumbuka kwamba unaweza kuanza maoni na `<!--` na kumaliza maoni na `-->`

Ongeza maoni juu ya HTML yako yanayosema `Code below this line should not be changed`

# --hints--

Unapaswa kuanza maoni na `<!--` juu ya HTML yako.

```js
assert.match(code,(/^\s*<!--/));
```

Maoni yako yanapaswa kuwa na maandishi `Code below this line should not be changed`.

```js
assert.match(code,(/<!--(?!(>|->|.*-->.*this line))\s*.*this line.*\s*-->/gi));
```

Unapaswa kufunga maoni yako na `-->`.

```js
assert.match(code,(/-->.*\n+.+/g));
```

Unapaswa kuwa na idadi sawa ya waanzishaji na wafunga maoni.

```js
assert.equal(code.match(/<!--/g).length,code.match(/-->/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<!-- Code below this line should not be changed -->
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
