---
id: bad87fee1348bd9aec908850
title: Tumia mtindo wa kitufe cha Bootstrap chaguomsingi
challengeType: 0
forumTopicId: 16657
dashedName: apply-the-default-bootstrap-button-style
---

# --description--

Bootstrap ina darasa la vitu la kitufe lingine linaloitwa `btn-default`.

Tumia darasa la vitu la `btn` na `btn-default` kwa kila kipengele chako cha `button`.

# --hints--

Unapaswa kutumia darasa la vitu la `btn` kwa kila kipengele chako cha `button`.

```js
assert.lengthOf(document.querySelectorAll('.btn'),6);
```

Unapaswa kutumia darasa la vitu la `btn-default` kwa kila kipengele chako cha `button`.

```js
assert.lengthOf(document.querySelectorAll('.btn-default'), 6);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">
        <button></button>
        <button></button>
        <button></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">
        <button></button>
        <button></button>
        <button></button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">
        <button class="btn btn-default"></button>
        <button class="btn btn-default"></button>
        <button class="btn btn-default"></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">
        <button class="btn btn-default"></button>
        <button class="btn btn-default"></button>
        <button class="btn btn-default"></button>
      </div>
    </div>
  </div>
</div>
```
