---
id: bad87fee1348bd9bec908846
title: Unda safu ya Bootstrap
challengeType: 0
forumTopicId: 16813
dashedName: create-a-bootstrap-row
---

# --description--

Sasa tutaunda safu ya Bootstrap kwa ajili ya vipengele vyetu vya ndani ya mstari.

Unda kipengele cha `div` chini ya lebo ya `h3`, ukiwa na darasa la vitu la `row`.

# --hints--

Unapaswa kuongeza kipengele cha `div` chini ya kipengele chako cha `h3`.

```js
assert.lengthOf(document.querySelectorAll('div'),2);
assert.lengthOf(document.querySelectorAll('div.row h3.text-primary'),0);
assert.lengthOf(document.querySelectorAll('div.row + h3.text-primary'),0)
assert.lengthOf(document.querySelectorAll('h3.text-primary + div.row'),1);
```

Kipengele chako cha `div` kinapaswa kuwa na darasa la vitu la `row`

```js
const newDiv = document.querySelectorAll('div')?.[1]; 
assert.isTrue(newDiv?.classList?.contains('row'));
```

`row div` yako inapaswa kupangwa ndani ya `container-fluid div`

```js
assert.lengthOf(document.querySelectorAll('div.container-fluid div.row'),1);
```

Kipengele chako cha `div` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g).length ,code.match(/<div/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>

</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row"></div>
</div>
```
