---
id: bad87fee1348bd9aec908853
title: Ongeza sifa za kipengele za id kwa vipengele vya Bootstrap
challengeType: 0
forumTopicId: 16639
dashedName: add-id-attributes-to-bootstrap-elements
---

# --description--

Kumbuka kwamba pamoja na sifa za darasa la vitu, unaweza kumpa kila kipengele chako sifa ya kipengele ya `id`.

Kila id lazima iwe ya kipekee kwa kipengele fulani na itumike mara moja tu kwa kila ukurasa.

Tumpaie id ya kipekee kila kipengele chetu cha `div` cha darasa la `well`.

Kumbuka kwamba unaweza kumpa kipengele id kama hii:

```html
<div class="well" id="center-well">
```

Mpatie well upande wa kushoto id ya `left-well`. Mpatie well upande wa kulia id ya `right-well`.

# --hints--

`well` yako ya kushoto inapaswa kuwa na id ya `left-well`.

```js
const column = document.querySelectorAll('.col-xs-6')[0];
const leftWall = column?.querySelectorAll(`:scope ${'#left-well'}`);
assert.lengthOf(leftWall,1)
```

`well` yako ya kulia inapaswa kuwa na id ya `right-well`.

```js
const column = document.querySelectorAll('.col-xs-6')[1];
const rightWall = column?.querySelectorAll(`:scope ${'#right-well'}`);
assert.lengthOf(rightWall,1)
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well" id="left-well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well" id="right-well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
  </div>
</div>
```
