---
id: bad87fee1348bd9aec908746
title: Weka ukurasa wetu ndani ya div ya kontena-fluid ya Bootstrap
challengeType: 0
forumTopicId: 18198
dashedName: house-our-page-within-a-bootstrap-container-fluid-div
---

# --description--

Sasa hakikisha maudhui yote kwenye ukurasa wako yanajibadilisha kulingana na kifaa cha mkononi.

Tupange ndani kipengele chako cha `h3` ndani ya kipengele cha `div` chenye darasa la vitu `container-fluid`.

# --hints--

Kipengele chako cha `div` kinapaswa kuwa na darasa la vitu `container-fluid`.

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container-fluid'));
```

Kila kipengele chako cha `div` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);

assert.equal(code.match(/<\/div>/g).length ,code.match(/<div/g).length);
```

Kipengele chako cha `h3` kinapaswa kupangwa ndani ya kipengele cha `div`.

```js
const divElement = document.querySelector('div');
const divChildren = divElement?.querySelectorAll(`:scope ${'h3'}`)
assert.lengthOf(divChildren,1);
```

# --seed--

## --seed-contents--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```

# --solutions--

```html
<div class="container-fluid">
    <h3 class="text-primary text-center">jQuery Playground</h3>
</div>
```
