---
id: bad87fee1348bd9aec908849
title: Ongeza vipengele ndani ya visima vyako vya Bootstrap
challengeType: 0
forumTopicId: 16636
dashedName: add-elements-within-your-bootstrap-wells
---

# --description--

Sasa tuko na vipengele vingi vya `div` ndani ya kila safu ya safu yetu. Hii ndiyo kina cha kina tunachohitaji. Sasa tunaweza kuongeza vipengele vyetu vya `button`.

Panga ndani vipengele vitatu vya `button` ndani ya kila kipengele chako cha `div` chenye jina la darasa la vitu `well`.

# --hints--

Vipengele vitatu vya `button` vinapaswa kupangwa ndani ya kila kipengele chako cha `div` chenye darasa la vitu `well`.

```js
const buttonOne = document.querySelectorAll('div.well')?.[0];
const buttonOneChildren = buttonOne?.querySelectorAll(`:scope ${'button'}`);
assert.lengthOf(buttonOneChildren,3);

const buttonTwo = document.querySelectorAll('div.well')?.[1]
const buttonTwoChildren = buttonTwo?.querySelectorAll(`:scope ${'button'}`);
assert.lengthOf(buttonTwoChildren,3);
```

Unapaswa kuwa na jumla ya vipengele 6 vya `button`.

```js
assert.lengthOf(document.querySelectorAll('button'), 6);
```

Vipengele vyote vya `button` vinapaswa kuwa na lebo za kufunga.

```js
assert.match(code,/<\/button>/g);
assert.match(code,/<button/g);
assert.equal(code.match(/<\/button>/g).length,code.match(/<button/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">



      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">



      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">
        <button></button>
        <button></button>
        <button></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">
        <button></button>
        <button></button>
        <button></button>
      </div>
    </div>
  </div>
</div>
```
