---
id: bad87fee1348bd9aec908846
title: Unda kichwa cha habari cha Bootstrap
challengeType: 0
forumTopicId: 16812
dashedName: create-a-bootstrap-headline
---

# --description--

Sasa tujenge kitu kutoka mwanzo ili kuzoea ujuzi wetu wa HTML, CSS na Bootstrap.

Tutajenga uwanja wa majaribio wa jQuery, ambao hivi karibuni tutautumia katika changamoto zetu za jQuery.

Kuanza, unda kipengee cha `h3`, chenye maandishi `jQuery Playground`.

Paka rangi kipengee chako cha `h3` kwa darasa la vitu la Bootstrap la `text-primary`, na ukiletee katikati kwa darasa la vitu la Bootstrap la `text-center`.

# --hints--

Unapaswa kuongeza kipengee cha `h3` kwenye ukurasa wako.

```js
assert.lengthOf(document.querySelectorAll('h3'),1);
```

Kipengee chako cha `h3` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code,/<\/h3>/g);
assert.match(code,/<h3/g);
assert.equal( code.match(/<\/h3>/g).length , code.match(/<h3/g).length);
```

Kipengee chako cha `h3` kinapaswa kupakwa rangi kwa kutumia darasa la vitu `text-primary`

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-primary'));
```

Kipengee chako cha `h3` kinapaswa kuletwa katikati kwa kutumia darasa la vitu `text-center`

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-center'));
```

Kipengee chako cha `h3` kinapaswa kuwa na maandishi `jQuery Playground`.

```js
assert.match(document.querySelector('h3')?.textContent, /jquery(\s)+playground/gi);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```
