---
id: bad87fee1348bd9aec908848
title: Unda visima vya Bootstrap
challengeType: 0
forumTopicId: 16825
dashedName: create-bootstrap-wells
---

# --description--

Bootstrap ina darasa la vitu linaloitwa `well` ambalo linaweza kuunda hisia ya kina kwa safu zako.

Panga ndani kipengele kimoja cha `div` chenye darasa la `well` ndani ya kila kipengele chako cha `col-xs-6` `div`.

# --hints--

Unapaswa kuongeza kipengele cha `div` chenye darasa la `well` ndani ya kila kipengele chako cha `div` chenye darasa la `col-xs-6`

```js
const wells = document.querySelectorAll('div.col-xs-6 > div.well');
assert.lengthOf( wells,2 ); 
```

Vipengele vyote viwili vya `div` vyenye darasa la `col-xs-6` vinapaswa kupangwa ndani ya kipengele chako cha `div` chenye darasa la `row`.

```js
assert.lengthOf(document.querySelectorAll('div.row > div.col-xs-6'),2);
```

Vipengele vyote vya `div` vinapaswa kuwa na lebo za kufunga.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g)?.length , code.match(/<div/g)?.length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">

    </div>
    <div class="col-xs-6">

    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well"></div>
    </div>
    <div class="col-xs-6">
      <div class="well"></div>
    </div>
  </div>
</div>
```
