---
id: 68da604b8b6012a33e2344fb
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذه الورشة، ستنفذ دالة تولد كل التركيبات الصحيحة من الأقواس العادية باستخدام طريقة البحث بالاتساع أولاً (BFS). على سبيل المثال، التركيبات الصحيحة لزوجين من الأقواس هي `(())` و `()()`.

ابدأ بإنشاء دالة باسم `gen_parentheses` تحتوي على معلمة واحدة هي `pairs`. في الوقت الحالي، أرجع قائمة فارغة من الدالة.

# --hints--

يجب أن تعرف دالة باسم `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("gen_parentheses")
`) })
```

يجب أن تحتوي دالتك `gen_parentheses` على معلمة واحدة باسم `pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_args("pairs")
`) })
```

يجب أن تُرجع دالتك قائمة فارغة.

```js
({ test: () => runPython(`
foo = _Node(_code).find_function("gen_parentheses")
assert foo.has_return("[]") or foo.has_return("list()")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
