---
id: 68da6138d70960c197ec5eed
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

بعد ذلك، يجب التحقق من أن عدد الأزواج لا يقل عن واحد، إذ لا يمكنك توليد تركيبات الأقواس العادية إذا كان عدد الأزواج صفرًا أو سالبًا.

أضف جملة `if` أخرى للتحقق مما إذا كان `pairs` أقل من `1`. إذا تحقق هذا الشرط، أرجع السلسلة النصية `The number of pairs should be at least 1`.

# --hints--

يجب أن تحتوي دالة `gen_parentheses` على جملة `if` ثانية.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[1]
`) })
```

يجب أن تتحقق جملة `if` الثانية مما إذا كان `pairs` أقل من `1`.

```js
({ test: () => runPython(`
cond = _Node(_code).find_function("gen_parentheses").find_ifs()[1].find_conditions()[0]
assert cond.is_equivalent("pairs < 1") or cond.is_equivalent("1 > pairs")
`) })
```

يجب أن ترجع السلسلة النصية `The number of pairs should be at least 1` عندما يكون `pairs` أقل من `1`.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be at least 1'
assert gen_parentheses(0) == msg and gen_parentheses(-1) == msg
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
--fcc-editable-region--
    
--fcc-editable-region--
    
    return []
```
