---
id: 68da6138d70960c197ec5ef3
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

الآن عليك التعامل مع إضافة الأقواس العادية المغلقة. القاعدة الأساسية هي أنه يمكنك إضافة قوس عادي مغلق فقط إذا حافظت على التوازن، أي يجب أن يكون هناك أقواس عادية مفتوحة أكثر مستخدمة من الأقواس المغلقة.

لا تزال داخل كتلة `else`، أضف جملة `if` ثانية للتحقق مما إذا كان `closes_used < opens_used`.

إذا تحقق هذا الشرط، أضف صف مرتب جديد آخر إلى `queue`: `(current + ')', opens_used, closes_used + 1)`. هذا يمثل الحالة بعد إضافة قوس عادي مغلق.

# --hints--

يجب أن يكون لديك جملة `if` ثانية داخل كتلة `else` الخاصة بك.

```js
({ test: () => runPython(`
import ast
assert isinstance(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1], ast.If)
`) })
```

يجب أن تتحقق جملة `if` المتسكنة الثانية مما إذا كان `closes_used < opens_used`.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1].test).is_equivalent("closes_used < opens_used")
`) })
```

يجب أن تضيف `(current + ')', opens_used, closes_used + 1)` إلى القائمة عندما يكون الشرط صحيحًا.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1]).has_call("queue.append((current + ')', opens_used, closes_used + 1))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        if len(current) == 2 * pairs:
            result.append(current)
        else:
            if opens_used < pairs:
                queue.append((current + '(', opens_used + 1, closes_used))
--fcc-editable-region--
            
--fcc-editable-region--

    return result

print(gen_parentheses(1))
```
