---
id: 68da6138d70960c197ec5ef0
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

الآن ستنفذ الحلقة الرئيسية لـ BFS. أنشئ حلقة `while` تستمر طالما أن `queue` ليست فارغة (طالما أن `queue` تُقيّم إلى `True`).

داخل الحلقة، اطبع `queue`.

# --hints--

يجب أن تحتوي دالتك `gen_parentheses` على حلقة `while`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0]
`) })
```

يجب أن تحتوي حلقة `while` على الشرط `queue`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_conditions()[0].is_equivalent("queue")
`) })
```

يجب أن تطبع `queue` داخل حلقة `while`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_call("print(queue)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    
--fcc-editable-region--

    return result
```
