---
id: 68da6138d70960c197ec5eec
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

قبل تنفيذ الخوارزمية الأساسية، يجب التحقق من صحة المُدخَل. يجب أن تكون معلمة `pairs` عددًا صحيحًا، لأنها تمثل عدد أزواج الأقواس العادية التي ستُولد.

أضف جملة `if` في بداية دالتك للتحقق مما إذا كانت `pairs` ليست عددًا صحيحًا. استخدم الدالة `isinstance()` لذلك.

إذا تحقق الشرط، أرجع السلسلة النصية `The number of pairs should be an integer`.

# --hints--

يجب أن تحتوي دالتك `gen_parentheses` على جملة `if`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0]
`) })
```

يجب أن تتحقق جملة `if` مما إذا كانت `pairs` ليست عددًا صحيحًا باستخدام `not isinstance(pairs, int)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(pairs, int)")
`) })
```

يجب أن تُرجع السلسلة النصية `The number of pairs should be an integer` عندما لا يكون المُدخَل عددًا صحيحًا.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be an integer'
args = [1.1, "1", [1, 2], {}]
assert  all(gen_parentheses(arg) == msg for arg in args)
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
--fcc-editable-region--
    
--fcc-editable-region--
    
    return []
```
