---
id: 68da6138d70960c197ec5eee
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

الآن ستُعد هيكل البيانات لتخزين نتائجك. أنشئ متغيرًا باسم `result` وابدأه كقائمة فارغة. ستخزن هذه القائمة كل تراكيب الأقواس العادية الصالحة التي تُولدها.

حدّث جملة الإرجاع لتُرجع `result` بدلًا من قائمة فارغة.

# --hints--

يجب إعلان متغير باسم `result` في دالة `gen_parentheses` الخاصة بك.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("result")
`) })
```

يجب أن تبدأ `result` كقائمة فارغة.

```js
({ test: () => runPython(`
_result = _Node(_code).find_function("gen_parentheses").find_variable("result")
assert _result.is_equivalent("result = []") or _result.is_equivalent("result = list()")
`) })
```

يجب أن تُرجع `result` في نهاية دالتك.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_return("result")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
--fcc-editable-region--
    
    return []
--fcc-editable-region--
```
