---
id: 68da6138d70960c197ec5ef1
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

الآن عليك تحديد متى تبني تركيبة كاملة من الأقواس العادية. التركيبة الكاملة طولها يساوي ضعف عدد الأزواج (لأن كل زوج يضم قوسًا افتتاحيًا وقوسًا إغلاقًا).

داخل حلقة `while`، أضف جملة `if` للتحقق مما إذا كان `len(current) == 2 * pairs`. عندما يكون هذا الشرط صحيحًا، أضف `current` إلى قائمة `result`.

# --hints--

يجب أن تحتوي حلقة `while` على جملة `if`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0]
`) })
```

يجب أن تتحقق جملة `if` مما إذا كان `len(current) == 2 * pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("len(current) == 2 * pairs")
`) })
```

يجب أن تضيف `current` إلى قائمة `result` داخل جملة `if` الجديدة.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("result.append(current)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
--fcc-editable-region--
        
--fcc-editable-region--

    return result

print(gen_parentheses(1))
```
