---
id: 68da6138d70960c197ec5eef
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

بالنسبة لطريقة البحث بالاتساع أولاً، ستستخدم قائمة انتظار لتتبع الحالات المختلفة أثناء بناء تركيبات الأقواس العادية. كل حالة ستُمثّل كصف مرتب يحتوي على ثلاثة عناصر:

- السلسلة النصية الحالية التي تُبنى
- عدد الأقواس العادية المفتوحة المستخدمة حتى الآن
- عدد الأقواس العادية المغلقة المستخدمة حتى الآن

أنشئ متغيرًا باسم `queue` وابدأ تهيئته بقائمة تحتوي على صف مرتب واحد: `('', 0, 0)`. هذا يمثل الحالة الابتدائية بسلسلة نصية فارغة وعدد أقواس صفر.

# --hints--

يجب إعلان متغير باسم `queue` في دالة `gen_parentheses` الخاصة بك.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("queue")
`) })
```

يجب تهيئة `queue` بقائمة تحتوي على الصف المرتب `('', 0, 0)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("queue").is_equivalent("queue = [('', 0, 0)]")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
--fcc-editable-region--
    
--fcc-editable-region--
    result = []
    return result
```
