---
id: 68da6138d70960c197ec5ef2
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

إذا لم تكتمل السلسلة النصية الحالية بعد، يجب استكشاف الحالات الممكنة التالية. أضف بند `else` إلى جملة `if` الخاصة بك.

داخل كتلة `else`، ستتعامل مع إضافة الأقواس العادية المفتوحة. أضف جملة `if` للتحقق مما إذا كان `opens_used < pairs`. هذا يضمن ألا تستخدم أكثر من عدد الأقواس المفتوحة المسموح به.

إذا تحقق هذا الشرط، أضف صفًا مرتبًا جديدًا إلى `queue`: `(current + '(', opens_used + 1, closes_used)`. هذا يمثل الحالة بعد إضافة قوس عادي مفتوح.

# --hints--

يجب أن يكون لديك بند `else` لجملة `if` داخل حلقة `while`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse
`) })
```

يجب أن يكون لديك جملة `if` داخل كتلة `else` تتحقق مما إذا كان `opens_used < pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1].is_equivalent("opens_used < pairs")
`) })
```

يجب أن تضيف `(current + '(', opens_used + 1, closes_used)` إلى قائمة الانتظار عندما يكون الشرط صحيحًا.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[1].has_call("queue.append((current + '(', opens_used + 1, closes_used))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        if len(current) == 2 * pairs:
            result.append(current)
--fcc-editable-region--
        
--fcc-editable-region--
    
    return result

print(gen_parentheses(1))
```
