---
id: 68e00c5349cbddfef6dbc053
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

داخل حلقة `while`، استخدم `queue.pop(0)` لإزالة العنصر الأول من الطابور والحصول عليه. هذا ينفذ سلوك الدخول أولاً خروج أولاً (FIFO) المميز لخوارزمية BFS.

فك هذا الصف المرتب إلى ثلاث متغيرات: `current`، و`opens_used`، و`closes_used`.

# --hints--

يجب أن تستخدم `queue.pop(0)` للحصول على العنصر الأول من الطابور.

```js
({ test: () => runPython(`
assert "queue.pop(0)" in str(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0])
`) })
```

يجب أن تفك النتيجة إلى ثلاث متغيرات: `current`، و`opens_used`، و`closes_used`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_stmt("current, opens_used, closes_used = queue.pop(0)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    while queue:
        print(queue)
--fcc-editable-region--
        
--fcc-editable-region--

    return result
```
