---
id: 587d8259367417b2b2512c85
title: Implementar el algoritmo de ordenación por selección
challengeType: 26
dashedName: implement-the-selection-sort-algorithm
---

# --description--

In this lab you will implement the selection sort algorithm. La ordenación por selección funciona seleccionando el valor mínimo de una lista e intercambiándolo con el primer valor de la misma. Después pasa a la segunda posición, busca y selecciona el valor más pequeño en los elementos restantes de la lista y lo mueve a la segunda posición. Continua realizando iteraciones en la lista y moviendo los elementos hasta que se llega a la última posición de la lista. Ahora la lista está ordenada. La ordenación por selección tiene una complejidad temporal cuadrática en todos los casos.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes tener una función `selectionSort`.
1. La función `selectionSort` debe tomar un arreglo y devolver un arreglo con los mismos elementos pero ordenados de menor a mayor.

# --hints--

`selectionSort` debería ser una función.

```js
assert.isFunction(selectionSort);
```

`selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` debe devolver un arreglo que no cambie excepto por el orden.

```js
assert.sameMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`selectionSort` debería devolver un arreglo ordenado (de menor a mayor).

```js
assert.sameOrderedMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
)
```

`selectionSort` no debe utilizar el método incorporado `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
