---
id: 63b607af6fcdb119aae9b16a
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

フォームでは、ユーザーがさまざまな食事タイプとそのカロリー数を追加できるようにしてください。

`fieldset` 要素を作成し、`id` を `breakfast` に設定してください。

その要素の中に、テキストが `legend` の `Breakfast` と、`div` が `class` に設定された空の `input-container` を作成してください。

# --hints--

`fieldset` の中に `form` 要素を作成してください。

```js
assert.exists(document.querySelector('form fieldset'));
```

`fieldset` 要素の後に `input` 要素が来るようにしてください。

```js
assert.equal(document.querySelector('form fieldset')?.previousElementSibling?.tagName, "INPUT");
```

`fieldset` 要素には `id` が `breakfast` に設定されている必要があります。

```js
assert.equal(document.querySelector('form fieldset')?.id, "breakfast");
```

`legend` の中に `fieldset` 要素を作成してください。

```js
assert.exists(document.querySelector('form fieldset legend'));
```

`legend` 要素にはテキスト `Breakfast` が含まれている必要があります。

```js
assert.equal(document.querySelector('form fieldset legend')?.innerText, "Breakfast");
```

`div` の中に `fieldset` 要素を作成してください。

```js
assert.exists(document.querySelector('form fieldset div'));
```

`div` 要素の後に `legend` 要素が来るようにしてください。

```js
assert.equal(document.querySelector('form fieldset div')?.previousElementSibling?.tagName, "LEGEND");
```

`div` 要素には `class` が `input-container` に設定されている必要があります。

```js
assert.equal(document.querySelector('form fieldset div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
