---
id: 614389f601bb4f611db98563
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

`.author` 要素の下に、クラス `div` を持つ新しい `social-icons` 要素を作成してください。

その新しい `a` の中に5つの `div` 要素を追加し、それぞれに以下の `href` 属性を設定してください。

- 最初の `a` 要素には `href` を `https://www.facebook.com/freecodecamp` に設定してください。
- 2番目の `a` 要素には `href` を `https://twitter.com/freecodecamp` に設定してください。
- 3番目の `a` 要素には `href` を `https://instagram.com/freecodecamp` に設定してください。
- 4番目の `a` 要素には `href` を `https://www.linkedin.com/school/free-code-camp` に設定してください。
- 5番目の `a` 要素には `href` を `https://www.youtube.com/freecodecamp` に設定してください。

# --hints--

新しい `div` 要素を作成してください。

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

新しい `div` 要素は `.author` 要素の後に配置してください。

```js
assert.equal(document.querySelector('.author')?.nextElementSibling?.localName, 'div');
```

新しい `div` 要素にはクラス `social-icons` を設定してください。

```js
assert.isTrue(document.querySelector('.author')?.nextElementSibling?.classList?.contains('social-icons'));
```

`.social-icons` 要素には5つの `a` 要素を含めてください。

```js
assert.lengthOf(document.querySelector('.social-icons')?.querySelectorAll('a'), 5);
```

最初の `a` 要素には `href` を `https://www.facebook.com/freecodecamp` に設定してください。

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[0]?.getAttribute('href'), 'https://www.facebook.com/freecodecamp');
```

2番目の `a` 要素には `href` を `https://twitter.com/freecodecamp` に設定してください。

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[1]?.getAttribute('href'), 'https://twitter.com/freecodecamp');
```

3番目の `a` 要素には `href` を `https://instagram.com/freecodecamp` に設定してください。

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[2]?.getAttribute('href'), 'https://instagram.com/freecodecamp');
```

4番目の `a` 要素には `href` を `https://www.linkedin.com/school/free-code-camp` に設定してください。

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[3]?.getAttribute('href'), 'https://www.linkedin.com/school/free-code-camp');
```

5番目の `a` 要素には `href` を `https://www.youtube.com/freecodecamp` に設定してください。

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[4]?.getAttribute('href'), 'https://www.youtube.com/freecodecamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
