---
id: 6143869bb45bd85e3b1926aa
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

`.heading` 要素の中に、`header` が `class` に設定された `hero` 要素を作成してください。

その要素の中に、`img` が `src` に、`https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png` が `alt` に、`freecodecamp logo` が `class` に設定された `hero-img` 要素を作成してください。

`loading` 要素の `img` 属性を `lazy` に設定すると、ブラウザーは画像リソースを必要になるまで（ユーザーが画像をスクロールして表示領域に入れたときなど）取得しません。さらに、遅延読み込みされた要素は非遅延要素が読み込まれるまで読み込まれないため、通信速度が遅いユーザーでも画像の読み込みを待たずにページの内容を閲覧できます。

新しい `img` 要素に `loading` 属性を `lazy` に設定してください。

# --hints--

`header` 要素を作成してください。

```js
assert.exists(document.querySelector('header'));
```

`header` 要素は `.heading` 要素の中にあるべきです。

```js
assert.equal(document.querySelector('header')?.parentElement?.className, 'heading');
```

`header` 要素は `class` が `hero` に設定されているべきです。

```js
assert.equal(document.querySelector('header')?.className, 'hero');
```

`img` 要素を作成してください。

```js
assert.exists(document.querySelector('img'));
```

`img` 要素は `header` 要素の中にあるべきです。

```js
assert.equal(document.querySelector('img')?.parentElement?.localName, 'header');
```

`img` 要素は `src` が `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png` に設定されているべきです。

```js
assert.equal(document.querySelector('img')?.getAttribute('src'), 'https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png');
```

`img` 要素は `alt` が `freecodecamp logo` に設定されているべきです。

```js
assert.equal(document.querySelector('img')?.getAttribute('alt'), 'freecodecamp logo');
```

`img` 要素は `loading` 属性が `lazy` に設定されているべきです。

```js
assert.equal(document.querySelector('img')?.getAttribute('loading'), 'lazy');
```

`img` 要素は `class` が `hero-img` に設定されているべきです。

```js
assert.equal(document.querySelector('img')?.className, 'hero-img');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
