---
id: 6143862a5e54455d262c212e
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

`body`の中に`main`要素を作成してください。その要素の中に、`section`に設定された`class`を持つ`heading`を作成してください。

# --hints--

`main` 要素があるはずです。

```js
assert.exists(document.querySelector('main'));
```

`main`要素は`body`要素の中にあるはずです。

```js
assert.equal(document.querySelector('main')?.parentElement?.localName, 'body');
```

`section`要素があるはずです。

```js
assert.exists(document.querySelector('section'));
```

`section` 要素は `main` 要素の中にある必要があります。

```js
assert.equal(document.querySelector('section')?.parentElement?.localName, 'main');
```

`section`要素は`class`に設定された`heading`を持つはずです。

```js
assert.isTrue(document.querySelector('section')?.classList?.contains('heading'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
