---
id: 614387cbefeeba5f3654a291
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

`header` 要素の後に、`div` が `class` に設定された `author` を作成してください。

その `div` の中に、`p` が `class` に設定された `author-name` 要素を作成し、テキストに `By freeCodeCamp` を入れてください。`freeCodeCamp` の部分は、`a` が `href` に設定され、`https://freecodecamp.org` が `target` に設定された `_blank` 要素で囲んでください。


その下に、クラスが `p` でテキストが `publish-date` の2つ目の `March 7, 2019` 要素を追加してください。

# --hints--

新しい `div` 要素を作成してください。

```js
assert.exists(document.querySelector('div'));
```

`div` 要素は `header` 要素の後に配置してください。

```js
assert.equal(document.querySelector('div')?.previousElementSibling?.localName, 'header');
```

`div` 要素は `class` が `author` に設定されている必要があります。

```js
assert.equal(document.querySelector('div')?.className, 'author');
```

新しい `p` 要素を2つ作成してください。

```js
assert.lengthOf(document.querySelectorAll('p'), 3);
```

新しい2つの `p` 要素は `div` 要素の中にある必要があります。

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

最初の新しい `p` 要素は `class` が `author-name` に設定されている必要があります。

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.className, 'author-name');
```

最初の新しい `p` 要素はテキストに `By freeCodeCamp` を持つ必要があります。

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.innerText, 'By freeCodeCamp');
```

2つ目の新しい `p` 要素は `class` が `publish-date` に設定されている必要があります。

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className, 'publish-date');
```

2つ目の新しい `p` 要素はテキストに `March 7, 2019` を持つ必要があります。

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText, 'March 7, 2019');
```

新しい `a` 要素を作成してください。

```js
assert.exists(document.querySelector('a'));
```

`a` 要素は最初の新しい `p` 要素の中にある必要があります。

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName, 'a');
```

`a` 要素は `href` が `https://freecodecamp.org` に設定されている必要があります。

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href'), 'https://freecodecamp.org');
```

`a` 要素は `target` が `_blank` に設定されている必要があります。

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target'), '_blank');
```

`a` 要素はテキスト `freeCodeCamp` を囲む必要があります。

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.innerText, 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
