---
id: 61438b5b66d76a6264430f2a
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

新しく作成したそれぞれの `a` 要素の中に `i` 要素を追加し、次のクラスを付けてください。

- 最初の `i` 要素にはクラス `fab fa-facebook-f` を付けてください
- 2番目の `i` 要素にはクラス `fab fa-twitter` を付けてください
- 3番目の `i` 要素にはクラス `fab fa-instagram` を付けてください
- 4番目の `i` 要素にはクラス `fab fa-linkedin-in` を付けてください
- 5番目の `i` 要素にはクラス `fab fa-youtube` を付けてください

# --hints--

`i` 要素が5つあるはずです。

```js
assert.lengthOf(document.querySelectorAll('i'), 5);
```

それぞれの `a` 要素は `i` 要素を1つだけ持つべきです。

```js
const aaaaa = document.querySelectorAll('.social-icons a');
assert.lengthOf(aaaaa, 5);
aaaaa?.forEach(
  a => {
    assert.lengthOf(a?.children, 1);
    assert.equal(a?.children?.[0]?.localName, 'i');
  }
);
```

それぞれの `i` 要素は `class` 属性を持ち、その中に `fab` を含むべきです。

```js
const iiiii = document.querySelectorAll('i');
assert.lengthOf(iiiii, 5);
iiiii?.forEach(
  i => {
    assert.isTrue(i?.classList?.contains('fab'));
  }
);
```

最初の `i` 要素は `class` 属性を持ち、その中に `fa-facebook-f` を含むべきです。

```js
assert.isTrue(document.querySelectorAll('i')?.[0]?.classList?.contains('fa-facebook-f'));
```

2番目の `i` 要素は `class` 属性を持ち、その中に `fa-twitter` を含むべきです。

```js
assert.isTrue(document.querySelectorAll('i')?.[1]?.classList?.contains('fa-twitter'));
```

3番目の `i` 要素は `class` 属性を持ち、その中に `fa-instagram` を含むべきです。

```js
assert.isTrue(document.querySelectorAll('i')?.[2]?.classList?.contains('fa-instagram'));
```

4番目の `i` 要素は `class` 属性を持ち、その中に `fa-linkedin-in` を含むべきです。

```js
assert.isTrue(document.querySelectorAll('i')?.[3]?.classList?.contains('fa-linkedin-in'));
```

5番目の `i` 要素は `class` 属性を持ち、その中に `fa-youtube` を含むべきです。

```js
assert.isTrue(document.querySelectorAll('i')?.[4]?.classList?.contains('fa-youtube'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--
        <div class="social-icons">
          <a href="https://www.facebook.com/freecodecamp/">
          </a>
          <a href="https://twitter.com/freecodecamp/">
          </a>
          <a href="https://instagram.com/freecodecamp">
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
          </a>
          <a href="https://www.youtube.com/freecodecamp">
          </a>
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
