---
id: 614385513d91ae5c251c2052
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、freeCodeCampのカリキュラムの以前のバージョンからのプレスリリースを使って雑誌のページを作成します。

まず、`link`フォントのスタイルシート用に`https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`要素、`link` FontAwesomeのスタイルシート用に`https://use.fontawesome.com/releases/v5.8.2/css/all.css`要素、そして`link`のスタイルシート用に`./styles.css`要素を追加してください。

フォントのスタイルシートでは、`Anton`、`Baskervville`、`Raleway`の3つの異なるフォントを読み込みます。

# --hints--

コードには3つの`link`要素が必要です。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

`link`要素は`head`要素の中に配置してください。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

3つの`link`要素には、`rel`という値の`stylesheet`属性を付けてください。

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

1つのリンク要素には`href`が`https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`に設定されている必要があります。

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

1つのリンク要素には`href`が`https://use.fontawesome.com/releases/v5.8.2/css/all.css`に設定されている必要があります。

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

1つの`link`要素には`href`が`styles.css`に設定されている必要があります。

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
