---
id: 6165d3b702a5d92ad970b30c
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

`img` 要素の後に、`h1` が `class` に設定され、テキストが `hero-title` に設定された `OUR NEW CURRICULUM` 要素を追加し、その後に `p` が `class` に設定され、テキストが `hero-subtitle` に設定された `Our efforts to restructure our curriculum with a more project-based focus` 要素を追加してください。

# --hints--

`h1` 要素を作成してください。

```js
assert.exists(document.querySelector('h1'));
```

`h1` 要素は `img` 要素の後に配置してください。

```js
assert.equal(document.querySelector('h1')?.previousElementSibling?.localName, 'img');
```

`h1` 要素は `class` を `hero-title` に設定してください。

```js
assert.equal(document.querySelector('h1')?.className, 'hero-title');
```

`h1` 要素はテキストを `OUR NEW CURRICULUM` に設定してください。

```js
assert.equal(document.querySelector('h1')?.textContent, 'OUR NEW CURRICULUM');
```

新しい `p` 要素を作成してください。

```js
assert.exists(document.querySelector('p'));
```

`p` 要素は `h1` 要素の後に来るべきです。

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.localName, 'h1');
```

`p` 要素は `class` を `hero-subtitle` に設定してください。

```js
assert.equal(document.querySelector('p')?.className, 'hero-subtitle');
```

`p` 要素はテキストを `Our efforts to restructure our curriculum with a more project-based focus` に設定してください。

```js
assert.equal(document.querySelector('p')?.textContent?.trim()?.replace(/\s{2,}/, ' '), 'Our efforts to restructure our curriculum with a more project-based focus');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
        </header>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
