---
id: 6619240f46cec8e04d77e03a
title: 基本函式練習 A 
challengeType: 1
dashedName: top-basic-functions-exercise-a
---

# --description--

創建一個接受整數的函式（程式）。如果整數小於 `10`，這個函式（程式）應該傳回給定的 `integer + 7`。如果整數大於或相等於 `10`，它應該傳回給定的 `integer - 3`。

函式（程式）的名稱應該是 `addOrSubtract`。

# --hints--

你應該有一個名為 `addOrSubtract` 的函式（程式）。

```js
assert.isFunction(addOrSubtract);
```

你的函式（程式）應該接受一個整數作為引數。

```js
assert.match(addOrSubtract.toString(), /\s*addOrSubtract\(\s*\w+\s*\)/);
```

如果整數小於 10，則你應該傳回該整數加 7。

```js
assert.strictEqual(addOrSubtract(5), 12);
```

如果整數大於或相等於 10，則你應該傳回該整數減 3。

```js
assert.strictEqual(addOrSubtract(10), 7);
```




# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addOrSubtract(num) {
  if (num < 10) {
    return num + 7;
  } else {
    return num - 3;
  }
}
```
