---
id: 661e17c6068359c3ccf2f4d8
title: 基本函式練習 D
challengeType: 1
dashedName: top-basic-functions-exercise-d
---

# --description--

撰寫一個名為 `lastLetter` 的函式（程式），該函式（程式）接受一個字串作為參數並傳回該字串的最後一個字母。

# --hints--

你應該有一個名為 `lastLetter` 的函式（程式）。

```js
assert.isFunction(lastLetter);
```

你的函式（程式）應該接受一個字串作為參數。 

```js
assert.match(lastLetter.toString(), /\s*lastLetter\(\s*\w+\s*\)/);
```

你應該傳回字串的最後一個字母。

```js
assert.strictEqual(lastLetter('Sem'), 'm');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function lastLetter(str) {
  return str[str.length - 1];
}
```
