---
id: 661e131f068359c3ccf2f4d6
title: 基本函式練習 B
challengeType: 1
dashedName: top-basic-functions-exercise-b
---

# --description--

撰寫一個名為 `multiply` 的函式（程式），該函式接受兩個參數並傳回它們的乘積。

# --hints--

你應該有一個名為 `multiply` 的函式（程式）。

```js
assert.isFunction(multiply);
```

你的函式（程式）應該接受兩個整數作為引數。

```js
assert.match(multiply.toString(), /\s*multiply\(\s*\w+\s*,\s*\w+\s*\)/);
```

你應該傳回兩個整數的乘積。

```js
assert.strictEqual(multiply(10, 10), 100);
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js 
function multiply(a, b) {
  return a * b;
}
```
