---
id: 661e151f068359c3ccf2f4d7
title: 基本函式練習 C
challengeType: 1
dashedName: top-basic-functions-exercise-c
---

# --description--

撰寫一個名為 `capitalize` 的函式（程式），該函式（程式）接受一個字串作為參數，並傳回一個首字母大寫的新字串。

# --hints--

你應該有一個名為 `capitalize` 的函式（程式）。

```js
assert.isFunction(capitalize);
```

你的函式（程式）應該接受一個字串作為參數。 

```js
assert.match(capitalize.toString(), /\s*capitalize\(\s*\w+\s*\)/);
```

你的函式應該傳回一個首字母大寫的新字串。

```js
assert.strictEqual(capitalize('sem'), 'Sem');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function capitalize(str) {
  return str[0].toUpperCase() + str.slice(1);
}
```
