---
id: 66fbcf750a62784cf11f5630
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

你現在應該將這三個物件推入 `recipes` 陣列。為此，你可以使用 `push()` 方法。

使用 `push()` 方法將所有食譜物件推入 `recipes` 陣列。請確保依序推入 `recipe1`、`recipe2` 和 `recipe3`。

同時刪除 `recipe1Name`、`recipe2CookingTime` 和 `recipe3Ingredients` 變數，以及記錄那些變數的 `console.log` 述語。

# --hints--

你應該將所有食譜物件推入 `recipes` 陣列。請確保依照它們宣告的順序推入。

```js
assert.deepEqual(recipes, [recipe1, recipe2, recipe3]);
```

你應該移除 `recipe1Name` 變數。

```js
assert.throws(() => recipe1Name, ReferenceError);
```

你應該移除 `console.log(recipe1Name);` 述語。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1Name\s*\)\s*;?/);
```

你應該移除 `recipe2CookingTime` 變數。

```js
assert.throws(() => recipe2CookingTime, ReferenceError);
```

你應該移除 `console.log(recipe2CookingTime);` 述語。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2CookingTime\s*\)\s*;?/);
```

你應該移除 `recipe3Ingredients` 變數。

```js
assert.throws(() => recipe3Ingredients, ReferenceError);
```

你應該移除 `console.log(recipe3Ingredients);` 述語。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe3Ingredients\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--
const recipe1Name = recipe1.name;
console.log(recipe1Name);

const recipe2CookingTime = recipe2.cookingTime;
console.log(recipe2CookingTime);

const recipe3Ingredients = recipe3.ingredients;
console.log(recipe3Ingredients);
--fcc-editable-region--
```
