---
id: 66fbcf750a62784cf11f5634
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

是時候測試每一個函式（程式）。你可以使用任何一個配方來進行測試，但這個教學將從 `recipe1` 開始。

創建兩個新的變數：`recipe1TotalIngredients` 和 `recipe1DifficultyLevel`。透過呼叫對應函式並傳入適當的 `recipe1` 屬性，為它們指定值。

最後，將每個變數記錄到主控台以查看結果。

# --hints--

你應該創建一個 `recipe1TotalIngredients` 變數。

```js
assert.isNotNull(recipe1TotalIngredients);
```

你的 `recipe1TotalIngredients` 變數應該設定為 `getTotalIngredients(recipe1.ingredients)`。

```js
assert.equal(recipe1TotalIngredients, getTotalIngredients(recipe1.ingredients));
```

你應該將 `recipe1TotalIngredients` 記錄到主控台。

```js
assert.match(code, /console\.log\(recipe1TotalIngredients\);?/);
```

你應該創建一個 `recipe1DifficultyLevel` 變數。

```js
assert.isNotNull(recipe1DifficultyLevel);
```

你的 `recipe1DifficultyLevel` 變數應該設定為 `getDifficultyLevel(recipe1.cookingTime)`。

```js
assert.equal(recipe1DifficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

你應該將 `recipe1DifficultyLevel` 記錄到主控台。

```js
assert.match(code, /console\.log\(recipe1DifficultyLevel\);?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
