---
id: 66fbcf750a62784cf11f5636
title: 步驟 12
challengeType: 1
dashedName: step-12
---

# --description--

對 `recipe2` 和 `recipe3` 的 `totalIngredients` 及 `difficultyLevel` 屬性重複該處理程序。

# --hints--

你應該存取 `recipe2` 的 `totalIngredients` 屬性。

```js
assert.isNotNull(recipe2.totalIngredients);
```

你應該將使用 `recipe2.ingredients` 呼叫 `getTotalIngredients` 的結果指定到 `recipe2` 的 `totalIngredients` 屬性。

```js
assert.strictEqual(recipe2.totalIngredients, getTotalIngredients(recipe2.ingredients));
```

你應該存取 `recipe2` 的 `difficultyLevel` 屬性。

```js
assert.isNotEmpty(recipe2.difficultyLevel);
```

你應該將呼叫 `getDifficultyLevel` 並以 `recipe2.cookingTime` 作為參數的結果指定給 `recipe2` 的 `difficultyLevel` 屬性。

```js
assert.strictEqual(recipe2.difficultyLevel, getDifficultyLevel(recipe2.cookingTime));
```

你應該存取 `recipe3` 的 `totalIngredients` 屬性。

```js
assert.isNotNull(recipe3.totalIngredients);
```

你應該將使用 `recipe3.ingredients` 呼叫 `getTotalIngredients` 的結果指定到 `recipe3` 的 `totalIngredients` 屬性。

```js
assert.strictEqual(recipe3.totalIngredients, getTotalIngredients(recipe3.ingredients));
```

你應該存取 `recipe3` 的 `difficultyLevel` 屬性。

```js
assert.isNotEmpty(recipe3.difficultyLevel);
```

你應該將使用 `recipe3.cookingTime` 呼叫 `getDifficultyLevel` 的結果指定到 `recipe3` 的 `difficultyLevel` 屬性。

```js
assert.strictEqual(recipe3.difficultyLevel, getDifficultyLevel(recipe3.cookingTime));
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

const recipe1TotalIngredients = getTotalIngredients(recipe1.ingredients);
console.log(recipe1TotalIngredients);

const recipe1DifficultyLevel = getDifficultyLevel(recipe1.cookingTime);
console.log(recipe1DifficultyLevel);

recipe1.totalIngredients = getTotalIngredients(recipe1.ingredients);
recipe1.difficultyLevel = getDifficultyLevel(recipe1.cookingTime);

--fcc-editable-region--

--fcc-editable-region--
```
