---
id: 66fbcf750a62784cf11f5633
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

創建一個 `getDifficultyLevel` 函式（程式），該函式（程式）接受一個表示烹飪時間的數字作為參數。

如果烹飪時間小於或相等於 `30`，函式應該傳回 `"easy"`。如果小於或相等於 `60`，函式應該傳回 `"medium"`。否則，函式應該傳回 `"hard"`。

# --hints--

你應該創建一個 `getDifficultyLevel` 函式（程式）。

```js
assert.isFunction(getDifficultyLevel);
```

你的 `getDifficultyLevel` 函式（程式）應該只有一個參數。

```js
assert.lengthOf(getDifficultyLevel, 1);
```

當烹飪時間小於或相等於 `30` 時，你的 `getDifficultyLevel` 函式（程式）應該傳回 `"easy"`。

```js
assert.strictEqual(getDifficultyLevel(10), "easy");
assert.strictEqual(getDifficultyLevel(20), "easy");
assert.strictEqual(getDifficultyLevel(29), "easy");
assert.strictEqual(getDifficultyLevel(30), "easy");
```

當烹飪時間大於 `31` 且小於或相等於 `60` 時，你的 `getDifficultyLevel` 函式（程式）應該傳回 `"medium"`。

```js
assert.strictEqual(getDifficultyLevel(31), "medium");
assert.strictEqual(getDifficultyLevel(40), "medium");
assert.strictEqual(getDifficultyLevel(50), "medium");
assert.strictEqual(getDifficultyLevel(60), "medium");
```

當烹飪時間大於 `60` 時，你的 `getDifficultyLevel` 函式（程式）應該傳回 `"hard"`。

```js
assert.strictEqual(getDifficultyLevel(61), "hard");
assert.strictEqual(getDifficultyLevel(65), "hard");
assert.strictEqual(getDifficultyLevel(70), "hard");
assert.strictEqual(getDifficultyLevel(75), "hard");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

--fcc-editable-region--

--fcc-editable-region--
```
