---
id: 66fbcf750a62784cf11f562d
title: 步驟 2
challengeType: 1
dashedName: step-2
---

# --description--

創建一個名為 `recipe1` 的物件。在 `recipe1` 物件內，創建一個 `name` 屬性，值為 `"Spaghetti Carbonara"`。

同樣在 `recipe1` 物件內，創建一個 `ingredients` 屬性，其值為一個陣列。該陣列應包含 `spaghetti`、`Parmesan cheese`、`pancetta` 和 `black pepper`。

# --hints--

你應該創建一個名為 `recipe1` 的物件。

```js
assert.isObject(recipe1);
```

你的 `recipe1` 物件應該有一個 `name` 屬性。

```js
assert.property(recipe1, 'name');
```

你的 `name` 屬性應該設定為 `Spaghetti Carbonara`。

```js
assert.equal(recipe1.name, 'Spaghetti Carbonara');
```

你的 `recipe1` 物件應該有一個 `ingredients` 屬性。

```js
assert.property(recipe1, 'ingredients');
```

你的 `ingredients` 屬性應該有一個陣列值。

```js
assert.isArray(recipe1.ingredients);
```

你的 `ingredients` 屬性的陣列值應該包含 `spaghetti`、`Parmesan cheese`、`pancetta` 和 `black pepper`。

```js
assert.deepEqual(recipe1.ingredients, ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];
--fcc-editable-region--

--fcc-editable-region--
```
