---
id: 66fbcf750a62784cf11f5635
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

你現在可以為 `recipes` 陣列中的每個項目填入 `totalIngredients` 和 `difficultyLevel` 屬性的值。

現在，存取 `recipe1` 的 `totalIngredients` 和 `difficultyLevel`，並將它們設定為函式呼叫和引數的適當結果。

# --hints--

你應該存取 `recipe1` 的 `totalIngredients` 屬性。

```js
assert.isNotNull(recipe1.totalIngredients);
```

你應該將使用 `recipe1.ingredients` 呼叫 `getTotalIngredients` 的結果指定到 `recipe1` 的 `totalIngredients` 屬性。

```js
assert.strictEqual(recipe1.totalIngredients, getTotalIngredients(recipe1.ingredients));
```

你應該存取 `recipe1` 的 `difficultyLevel` 屬性。

```js
assert.isNotEmpty(recipe1.difficultyLevel);
```

你應該將呼叫 `getDifficultyLevel` 並以 `recipe1.cookingTime` 作為參數的結果指定給 `recipe1` 的 `cookingTime` 屬性。

```js
assert.strictEqual(recipe1.difficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

const recipe1TotalIngredients = getTotalIngredients(recipe1.ingredients);
console.log(recipe1TotalIngredients);

const recipe1DifficultyLevel = getDifficultyLevel(recipe1.cookingTime);
console.log(recipe1DifficultyLevel);

--fcc-editable-region--

--fcc-editable-region--
```
