---
id: 66fbcdfed6848e48947dbbc9
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將使用 JavaScript 物件創建一個食譜追蹤器。

首先建立一個名為 `recipes` 的空陣列。這是你稍後會將食譜物件推入的陣列。

# --hints--

你應該創建一個名為 `recipes` 的陣列。

```js
assert.isArray(recipes);
```

你的 `recipes` 陣列應該是空的。

```js
assert.isEmpty(recipes);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
