---
id: 66fbcf750a62784cf11f5632
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

創建一個 `getTotalIngredients` 函式（程式），該函式接受一個引數，表述一個包含材料的陣列，並傳回從傳入函式的陣列中材料的數量。

# --hints--

你應該創建一個 `getTotalIngredients` 函式（程式）。

```js
assert.isFunction(getTotalIngredients);
```

你的 `getTotalIngredients` 函式（程式）應該只有一個參數。

```js
assert.lengthOf(getTotalIngredients, 1);
```

你的 `getTotalIngredients` 函式應該傳回一個數字。

```js
assert.isNumber(getTotalIngredients(recipe1.ingredients));
```

`getTotalIngredients(recipe1.ingredients)` 應該傳回 `4`。

```js
assert.strictEqual(getTotalIngredients(recipe1.ingredients), 4);
```

`getTotalIngredients(recipe2.ingredients)` 應該傳回 `5`。

```js
assert.strictEqual(getTotalIngredients(recipe2.ingredients), 5);
```

`getTotalIngredients(recipe3.ingredients)` 應該傳回 `3`。

```js
assert.strictEqual(getTotalIngredients(recipe3.ingredients), 3);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

--fcc-editable-region--

--fcc-editable-region--
```
