---
id: 66fbcf750a62784cf11f562e
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

創建一個具有以下屬性和值的 `recipe2` 物件：

| Key | Value   |
| ----------- | ------- |
| `name`       | `Chicken Curry` |
| `ingredients`       | `["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]` |
| `cookingTime`       | `42` |
| `totalIngredients`   | `null`   |
| `difficultyLevel`   | `""`   |

# --hints--

你應該創建一個名為 `recipe2` 的物件。

```js
assert.isObject(recipe2);
```

你的 `recipe2` 物件應該有一個 `name` 屬性。

```js
assert.property(recipe2, "name");
```

你的 `name` 屬性應該設為 `Chicken Curry`。

```js
assert.equal(recipe2.name, "Chicken Curry");
```

你的 `recipe2` 物件應該有一個 `ingredients` 屬性。

```js
assert.property(recipe2, "ingredients");
```

你的 `ingredients` 屬性應該有一個陣列值。

```js
assert.isArray(recipe2.ingredients);
```

你的 `ingredients` 屬性的陣列值應該包含 `chicken breast`、`coconut milk`、`curry powder`、`onion` 和 `garlic`。

```js
assert.deepEqual(recipe2.ingredients, ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]);
```

你的 `recipe2` 物件應該有一個 `cookingTime` 屬性。

```js
assert.property(recipe2, "cookingTime");
```

你的 `cookingTime` 屬性值應該是數字。

```js
assert.isNumber(recipe2.cookingTime);
```

你的 `cookingTime` 屬性應該設定為 `42`。

```js
assert.equal(recipe2.cookingTime, 42);
```

你的 `recipe2` 物件應該有一個 `totalIngredients` 屬性。

```js
assert.property(recipe2, "totalIngredients");
```

你的 `totalIngredients` 屬性應該設為 `null`。

```js
assert.isNull(recipe2.totalIngredients);
```

你的 `recipe2` 物件應該有一個 `difficultyLevel` 屬性。

```js
assert.property(recipe2, "difficultyLevel");
```

你的 `difficultyLevel` 屬性應該設為空字串。

```js
assert.deepEqual(recipe2.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
  };

--fcc-editable-region--

--fcc-editable-region--
```
