---
id: 66ff02792b3042a2b811f475
title: 步驟 3
challengeType: 1
dashedName: step-3
---

# --description--

將以下屬性添加到 `recipe1` 物件：

| Key | Value   |
| ----------- | ------- |
|`cookingTime`|`22`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

屬性中值為 `null` 和空字串的部分會在你計算後再更新。

# --hints--

你的 `recipe1` 物件應該有一個 `cookingTime` 屬性。

```js
assert.property(recipe1, "cookingTime");
```

你的 `cookingTime` 屬性應該設定為 `22`。

```js
assert.equal(recipe1.cookingTime, 22);
```

你的 `recipe1` 物件應該有一個 `totalIngredients` 屬性。

```js
assert.property(recipe1, "totalIngredients");
```

你的 `totalIngredients` 屬性應該設為 `null`。

```js
assert.isNull(recipe1.totalIngredients);
```

你的 `recipe1` 物件應該有一個 `difficultyLevel` 屬性。

```js
assert.property(recipe1, "difficultyLevel");
```

你的 `difficultyLevel` 屬性應該設為空字串。

```js
assert.deepEqual(recipe1.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
--fcc-editable-region--
  
--fcc-editable-region--
};
```
