---
id: 670e4f45f7116c0f216a5177
title: 步驟 6
challengeType: 1
dashedName: step-6
---

# --description--

在你繼續之前，你應該練習如何從物件中存取屬性。

你可以使用點號（`.`）或括號（`[]`）表示法來執行此操作。以下是一個範例：

```js
const person = {
  name: "John",
  age: 30,
  job: "Software Engineer"
};

console.log(person.name); // John
console.log(person['age']);  // 30
```

存取 `recipe1` 的 `name` 屬性，並將它指定給變數 `recipe1Name`。

接著，存取 `recipe2` 的 `cookingTime` 屬性並將它指定給變數 `recipe2CookingTime`。

最後，存取 `recipe3` 的 `ingredients` 屬性並將其指定給變數 `recipe3Ingredients`。

確保你創建的所有變數都已記錄到主控台。

# --hints--

你應該創建一個 `recipe1Name` 變數。

```js
assert.isDefined(recipe1Name);
```

你應該將 `recipe1` 的 `name` 屬性值指定給你的 `recipe1name` 變數。

```js
assert.strictEqual(recipe1Name, recipe1.name);
```

你應該將 `recipe1Name` 記錄到主控台。

```js
assert.match(code, /console\.log\(\s*recipe1Name\s*\)/);
```

你應該創建一個 `recipe2CookingTime` 變數。

```js
assert.isDefined(recipe2CookingTime);
```

你應該將 `recipe2` 的 `cookingTime` 屬性值指定給你的 `recipe2CookingTime` 變數。

```js
assert.strictEqual(recipe2CookingTime, recipe2.cookingTime);
```

你應該將 `recipe2CookingTime` 記錄到主控台。

```js
assert.match(code, /console\.log\(recipe2CookingTime\)/);
```

你應該創建一個 `recipe3Ingredients` 變數。

```js
assert.isDefined(recipe3Ingredients);
```

你應該將 `recipe3` 的 `ingredients` 屬性值指定給你的 `recipe3Ingredients` 變數。

```js
assert.strictEqual(recipe3Ingredients, recipe3.ingredients);
```

你應該將 `recipe3Ingredients` 記錄到主控台。

```js
assert.match(code, /console\.log\(recipe3Ingredients\)/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
