---
id: 69161f4c0c9c8c48b3c843c5
title: 步驟 5
challengeType: 1
dashedName: step-5
---

# --description--

創建一個具有以下屬性和值的 `recipe3` 物件：

|Key|Value|
|-----------|-------|
|`name`|`Vegetable Stir Fry`|
|`ingredients`|`["broccoli", "carrot", "bell pepper"]`|
|`cookingTime`|`15`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

# --hints--

你應該創建一個名為 `recipe3` 的物件。

```js
assert.isObject(recipe3);
```

你的 `recipe3` 物件應該有一個 `name` 屬性。

```js
assert.property(recipe3, "name");
```

你的 `name` 屬性應該設定為 `Vegetable Stir Fry`。

```js
assert.equal(recipe3.name, "Vegetable Stir Fry");
```

你的 `recipe3` 物件應該有一個 `ingredients` 屬性。

```js
assert.property(recipe3, "ingredients");
```

你的 `ingredients` 屬性應該有一個陣列值。

```js
assert.isArray(recipe3.ingredients);
```

你的 `ingredients` 屬性的陣列值應該包含 `broccoli`、`carrot` 和 `bell pepper`。

```js
assert.deepEqual(recipe3.ingredients, ["broccoli", "carrot", "bell pepper"]);
```

你的 `recipe3` 物件應該有一個 `cookingTime` 屬性。

```js
assert.property(recipe3, "cookingTime");
```

你的 `cookingTime` 屬性值應該是數字。

```js
assert.isNumber(recipe3.cookingTime);
```

你的 `cookingTime` 屬性應該設定為 `15`。

```js
assert.equal(recipe3.cookingTime, 15);
```

你的 `recipe3` 物件應該有一個 `totalIngredients` 屬性。

```js
assert.property(recipe3, "totalIngredients");
```

你的 `totalIngredients` 屬性應該設為 `null`。

```js
assert.isNull(recipe3.totalIngredients);
```

你的 `recipe3` 物件應該有一個 `difficultyLevel` 屬性。

```js
assert.property(recipe3, "difficultyLevel");
```

你的 `difficultyLevel` 屬性應該設為空字串。

```js
assert.deepEqual(recipe3.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
