---
id: 66c8ba975ee7230e29f6c4ad
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Щоб перевірити, на який кредит користувач має право на основі `annualIncome` та `creditScore`, потрібно використати інструкцію `if/else if` або тернарний оператор безпосередньо у функції `getLoanMessage`. Потім у блоці кожної умови поверніть відповідне повідомлення.

Починаючи з кредиту на дуплекс, перевірте, чи `annualIncome` більший або дорівнює `minIncomeForDuplex` І чи `creditScore` більший або дорівнює `minCreditScoreForDuplex`.

Якщо ця умова істинна, тоді заявник має право на кредит на дуплекс і на інші кредити. Тож всередині перевірки поверніть рядок `"You qualify for a duplex, condo, and car loan."`

# --hints--

Ваша функція `getLoanMessage` має повертати рядок.

```js
assert.isString(getLoanMessage(65000, 750));
assert.isString(getLoanMessage(60000, 750));
assert.isString(getLoanMessage(65000, 700));
assert.isString(getLoanMessage(60000, 700));
```

Ваша функція `getLoanMessage` має повертати рядок `"You qualify for a duplex, condo, and car loan."`.

```js
assert.strictEqual(getLoanMessage(65000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(65000, 700), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 700), "You qualify for a duplex, condo, and car loan.");
```

Ваша функція `getLoanMessage` має повертати `undefined`, якщо річний дохід заявника та кредитний рейтинг не відповідають вимогам для кредиту на дуплекс.

```js
assert.isUndefined(getLoanMessage(59000, 690));
assert.isUndefined(getLoanMessage(59000, 700));
assert.isUndefined(getLoanMessage(59000, 750));
assert.isUndefined(getLoanMessage(60000, 690));
assert.isUndefined(getLoanMessage(65000, 690));
```
 
# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
