---
id: 66c8ba975ee7230e29f6c4af
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Якщо річний дохід заявника більший або дорівнює `minIncomeForCondo`, І якщо їх кредитний рейтинг більший або дорівнює `minCreditScoreForCondo`, тоді вони підходять для отримання кредиту на кондомініум і автомобіль.

Перевірте, чи це правда у функції `getLoanMessage`. Якщо так, поверніть рядок `"You qualify for a condo and car loan."`

# --hints--

Ваша функція `getLoanMessage` має повертати рядок, якщо заявник підходить для кредиту на кондомініум і автомобіль.

```js
assert.isString(getLoanMessage(45000, 680));
```

Ваша функція `getLoanMessage` має повертати рядок `"You qualify for a condo and car loan."`.

```js
assert.strictEqual(getLoanMessage(45000, 680), "You qualify for a condo and car loan.");
```

Ваша функція `getLoanMessage` має повертати `undefined`, якщо річний дохід і кредитний рейтинг заявника не відповідають вимогам для кредиту на кондомініум.

```js
assert.strictEqual(getLoanMessage(45000, 650), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
