---
id: 66c8ba975ee7230e29f6c4b0
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Тепер слід перевірити, чи кандидат підходить лише для автокредиту. Якщо так, поверніть рядок `"You qualify for a car loan."`.

# --hints--

Ваша функція `getLoanMessage` має повертати рядок, якщо кандидат підходить для автокредиту.

```js
assert.isString(getLoanMessage(30000, 650));
```

Ваша функція `getLoanMessage` має повертати рядок `"You qualify for a car loan."`.

```js
assert.strictEqual(getLoanMessage(30000, 650), "You qualify for a car loan.");
```

Ваша функція `getLoanMessage` має повертати `undefined`, якщо річний дохід і кредитний рейтинг кандидата не відповідають вимогам для автокредиту.

```js
assert.strictEqual(getLoanMessage(30000, 600), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
--fcc-editable-region--
  } 
  
--fcc-editable-region--
}
```
