---
id: 66c8ba975ee7230e29f6c4b2
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Тепер настав час протестувати вашу функцію `getLoanMessage`.

Використайте таблицю нижче, щоб створити змінні `4` та їхні значення:

| Назва змінної | Значення   |
| ----------- | ------- |
| `duplexLoanMsg`       | `getLoanMessage(85000, 850)` |
| `condoLoanMsg`   | `getLoanMessage(65000, 690)`   |
| `carLoanMsg`   | `getLoanMessage(45000, 660)`   |
| `noLoanMsg`   | `getLoanMessage(25000, 550)`   |

Після цього виведіть кожну змінну в консоль, щоб побачити повідомлення.

Отже, ваш проєкт перевірки кваліфікації на позику завершено!

# --hints--

У вас має бути змінна з назвою `duplexLoanMsg`.

```js
assert.isNotNull(duplexLoanMsg);
```

Ваша змінна `duplexLoanMsg` має бути встановлена як результат виклику `getLoanMessage(85000, 850)`.

```js
assert.equal(duplexLoanMsg, getLoanMessage(85000, 850));
```

Ви маєте вивести змінну `duplexLoanMsg` у консоль.

```js
assert.match(code, /console\.log\(duplexLoanMsg\)/)
```

У вас має бути змінна з назвою `condoLoanMsg`.

```js
assert.isNotNull(condoLoanMsg);
```

Ваша змінна `condoLoanMsg` має бути встановлена як результат виклику `getLoanMessage(65000, 690)`.

```js
assert.equal(condoLoanMsg, getLoanMessage(65000, 690));
```

Ви маєте вивести змінну `condoLoanMsg` у консоль.

```js
assert.match(code, /console\.log\(condoLoanMsg\)/)
```

У вас має бути змінна з назвою `carLoanMsg`.

```js
assert.isNotNull(carLoanMsg);
```

Ваша змінна `carLoanMsg` має бути встановлена як результат виклику `getLoanMessage(45000, 660)`.

```js
assert.equal(carLoanMsg, getLoanMessage(45000, 660));
```

Ви маєте вивести змінну `carLoanMsg` у консоль.

```js
assert.match(code, /console\.log\(carLoanMsg\)/)
```

У вас має бути змінна з назвою `noLoanMsg`.

```js
assert.isNotNull(noLoanMsg);
```

Ваша змінна `noLoanMsg` має бути встановлена як результат виклику `getLoanMessage(25000, 550)`.

```js
assert.equal(noLoanMsg, getLoanMessage(25000, 550));
```

Ви маєте вивести змінну `noLoanMsg` у консоль.

```js
assert.match(code, /console\.log\(noLoanMsg\)/)
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

const duplexLoanMsg = getLoanMessage(85000, 850)
console.log(duplexLoanMsg)

const condoLoanMsg = getLoanMessage(65000, 690)
console.log(condoLoanMsg)

const carLoanMsg = getLoanMessage(45000, 660)
console.log(carLoanMsg)

const noLoanMsg = getLoanMessage(25000, 550)
console.log(noLoanMsg)
```
