---
id: 66c8ba975ee7230e29f6c4ac
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Коли користувач має право на отримання позики, ви захочете показати йому повідомлення в консолі.

Для цього створіть функцію, всередині якої будуть перевірки, що повертають, на яку позику претендує заявник.

Створіть порожню `getLoanMessage` функцію з параметрами `annualIncome` та `creditScore`.

# --hints--

Ви повинні створити `getLoanMessage` функцію

```js
assert.isFunction(getLoanMessage)
```

Ваша `getLoanMessage` функція має мати параметри `annualIncome` та `creditScore`.

```js
assert.match(getLoanMessage?.toString(), /\(\s*annualIncome,\s*creditScore\s*\)|\(\s*creditScore,\s*annualIncome\s*\)/);
```

`getLoanMessage` має бути порожньою функцією.

```js
assert.match(getLoanMessage?.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

--fcc-editable-region--

--fcc-editable-region--
```
