---
id: 66c8ba41a77db20a93f9d7a1
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви повторите умовні оператори JavaScript, створюючи застосунок для перевірки відповідності позики.

Застосунок перевірятиме, чи користувач підходить для отримання позики на дуплекс, автомобіль або кондомініум, базуючись на його річному доході та кредитному рейтингу. Кредитний рейтинг — це число, яке представляє, наскільки добре ви керуєте позиченими грошима.

Щоб почати, створіть наступні змінні та значення.

| Назва змінної | Значення   |
| ----------- | ------- |
| `minIncomeForDuplex`       | `60000` |
| `minCreditScoreForDuplex`   | `700`   |
| `minIncomeForCondo`   | `45000`   |
| `minCreditScoreForCondo`    | `680`   |
| `minIncomeForCar`   | `30000`   |
| `minCreditScoreForCar`   | `650`   |

# --hints--

Ви повинні створити змінну `minIncomeForDuplex`, встановлену в `60000`.

```js
assert.equal(minIncomeForDuplex, 60000);
```

Ви повинні створити `minCreditScoreForDuplex`, встановлену в `700`.

```js
assert.equal(minCreditScoreForDuplex, 700);
```

Ви повинні створити змінну `minIncomeForCar`, встановлену в `30000`.

```js
assert.equal(minIncomeForCar, 30000);
```

Ви повинні створити змінну `minCreditScoreForCar`, встановлену в `650`.

```js
assert.equal(minCreditScoreForCar, 650);
```

Ви повинні створити змінну `minIncomeForCondo`, встановлену в `45000`.

```js
assert.equal(minIncomeForCondo, 45000);
```

Ви повинні створити змінну `minCreditScoreForCondo`, встановлену в `680`.

```js
assert.equal(minCreditScoreForCondo, 680);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
