---
id: 66c8ba975ee7230e29f6c4b1
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Якщо річний дохід заявника та кредитний рейтинг нижчі за `minIncomeForCar` і `minCreditScoreForCar`, тоді він не підходить для жодного кредиту. Тож поверніть рядок `"You don't qualify for any loans."`

# --hints--

Ваша функція `getLoanMessage` має повертати рядок, якщо заявник не підходить для жодного кредиту.

```js
assert.isString(getLoanMessage(25000, 550));
```

Ваша функція `getLoanMessage` має повертати рядок `"You don't qualify for any loans."`, якщо заявник не відповідає вимогам жодного кредиту.

```js
assert.strictEqual(getLoanMessage(15000, 600), "You don't qualify for any loans.");
assert.strictEqual(getLoanMessage(25000, 550), "You don't qualify for any loans.");
assert.strictEqual(getLoanMessage(20000, 500), "You don't qualify for any loans.");
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
