---
id: 6716249b5405164036fd0b0d
title: Construir um Visualizador de Ordenação
challengeType: 25
dashedName: build-a-sorting-visualizer
demoType: onClick
---

# --description--

O algoritmo **Bubble Sort** ordena uma sequência de inteiros comparando pares de elementos adjacentes começando do início da sequência.Se o primeiro elemento for maior que o segundo, ele os troca.Então, ele prossegue com o seguinte par.Quando o último elemento da sequência é alcançado, inicia um novo ciclo a partir do início da sequência e repete o processo até que os elementos estejam ordenados.O algoritmo para após um ciclo ser concluído sem trocas.

Para este laboratório, você recebeu todo o HTML e CSS. Você usará JavaScript para completar o Bubble Sort Visualizer para que ele visualize cada passo necessário pelo algoritmo Bubble Sort para ordenar um array de cinco inteiros.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter uma função chamada `generateElement` que retorna um número inteiro aleatório entre `1` e `100`, inclusive.
1. Você deve ter uma função chamada `generateArray` que usa a função `generateElement` para retornar um array contendo cinco inteiros aleatórios.
1. Você deve ter uma função chamada `generateContainer` que cria e retorna um elemento `div` vazio.
1. Você deve ter uma função chamada `fillArrContainer` que recebe um elemento HTML como primeiro argumento e um array como segundo argumento.
1. O `fillArrContainer` deve receber um elemento como seu primeiro parâmetro e um array de inteiros como seu segundo parâmetro, então preencher o elemento com cinco elementos `span`, com cada `span` mostrando um dos inteiros do array.
1. Você deve ter uma função chamada `isOrdered` que recebe dois inteiros e retorna um booleano indicando se o primeiro inteiro é menor ou igual ao segundo.
1. Você deve ter uma função chamada `swapElements` que recebe um array de inteiros e um índice numérico.
1. A função `swapElements` deve modificar o array no local, trocando o elemento no índice passado pelo elemento seguinte se `isOrdered` retornar `false`.
1. Você deve ter uma função chamada `highlightCurrentEls` que recebe um elemento HTML e um índice numérico.
1. A função `highlightCurrentEls` deve definir a `border` do filho do elemento fornecido no índice dado, e do filho imediatamente após o índice, para ter um estilo `dashed`, uma cor `red` e uma largura de sua escolha.
1. Quando você clicar em `#generate-btn`, deve usar a função `fillArrContainer` para preencher `#starting-array` com cinco elementos `span`, cada um com um número aleatório como texto. Se presentes, outros elementos devem ser removidos de `#array-container`.
1. Você deve implementar o algoritmo Bubble Sort para que, após clicar em `#sort-btn`, `#array-container` contenha um elemento `div` para cada um dos passos necessários pelo algoritmo Bubble Sort para ordenar o array inicial, incluindo a `div` que representa o array inicial e uma `div` que representa o array ordenado. As funções que você criou até agora podem ser úteis aqui.
1. Cada `div` deve conter cinco elementos `span`, representando o array em seu estado atual de ordenação.
1. Após clicar em `#sort-btn`, `#starting-array` deve representar o passo inicial com o array inicial e os dois primeiros inteiros destacados.
1. Para cada etapa de ordenação, você deve usar `highlightCurrentEls` para destacar os dois números que estão sendo comparados e trocá-los na próxima etapa usando `swapElements`.

# --hints--

Você deve ter uma função chamada `generateElement`.

```js
assert.isFunction(generateElement);
```

Sua função `generateElement` deve retornar um número inteiro aleatório entre `1` e `100`, inclusive.

```js
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    assert.strictEqual(generateElement(), 26);
    assert.strictEqual(generateElement(), 9);
} finally {
    randomMocker.restore();
}
```

Você deve ter uma função chamada `generateArray`.

```js
assert.isFunction(generateArray)
```

Sua função `generateArray` deve fazer uso da função `generateElement`.

```js
let flag = false;
const temp = generateElement;
generateElement = () => flag = true;
try {
    generateArray();
    assert.isTrue(flag);
} finally {
    generateElement = temp;
}
```

Sua função `generateArray` deve retornar um array contendo cinco números inteiros aleatórios entre `1` e `100`.

```js
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {
    assert.deepEqual(generateArray(), [26, 9, 58, 23, 38])
} finally {
    randomMocker.restore();
}
```

Você deve ter uma função chamada `generateContainer`.

```js
assert.isFunction(generateContainer);
```

Sua função `generateContainer` deve retornar um elemento `div` vazio.

```js
const div = generateContainer();
assert.equal(div.tagName, "DIV");
assert.isEmpty(div.children);
```

Você deve ter uma função chamada `fillArrContainer`.

```js
assert.isFunction(fillArrContainer);
```

Seu `fillArrContainer()` deve receber um elemento como seu primeiro parâmetro e um array de inteiros como seu segundo parâmetro, então preencher o elemento com cinco elementos `span`, com cada `span` mostrando um dos inteiros do array.

```js
const testDiv = document.createElement("div");
const testArr = [15, 98, 17, 5, 63]
fillArrContainer(testDiv, testArr);
const children = testDiv.children;
assert.lengthOf(children, 5);
Array.from(children).forEach((el, i) => {
    assert.equal(el.tagName, "SPAN");
    assert.equal(el.innerText.trim(), testArr[i])
})
```

Você deve ter uma função chamada `isOrdered`.

```js
assert.isFunction(isOrdered);
```

Sua função `isOrdered` deve receber dois inteiros e retornar um booleano indicando se o primeiro inteiro é menor ou igual ao segundo.

```js
assert.lengthOf(isOrdered, 2);
assert.isTrue(isOrdered(2, 60));
assert.isFalse(isOrdered(10, 3));
assert.isTrue(isOrdered(5, 5));
```

Você deve ter uma função chamada `swapElements`.

```js
assert.isFunction(swapElements);
```

Sua função `swapElements` recebe um array de inteiros e um índice numérico como argumentos.Deve modificar o array passado no local trocando o elemento no índice fornecido com o elemento seguinte se o primeiro elemento for maior que o segundo.

```js
const testArr = [22, 4, 87, 47, 33];
swapElements(testArr, 0);
assert.deepEqual(testArr, [4, 22, 87, 47, 33]);
swapElements(testArr, 1);
assert.deepEqual(testArr, [4, 22, 87, 47, 33]);
swapElements(testArr, 2);
assert.deepEqual(testArr, [4, 22, 47, 87, 33]);
swapElements(testArr, 3);
assert.deepEqual(testArr, [4, 22, 47, 33, 87]);
```

Você deve ter uma função chamada `highlightCurrentEls`.

```js
assert.isFunction(highlightCurrentEls);
```

Sua função `highlightCurrentEls` deve dar aos descendentes do elemento especificado, localizados no índice fornecido e no próximo índice, uma borda que seja `tracejada`, `vermelha` e definida com uma largura de sua escolha.

```js
const testDiv = document.createElement("div");
document.querySelector("body").appendChild(testDiv)
for (let i = 0; i < 5; i++) {
    testDiv.appendChild(document.createElement("span"));
}
const redBorderRegex = /dashed (rgb\(255,\s*0,\s*0\)|#FF0000|#F00|hsl\(0,\s*100%,\s*50%\))/;
const revertBorder = () => {
    for (const el of children) {
        el.style.border = "revert";
    }
}
const children = testDiv.children;

for (let i = 0; i < 3; i++) {
    highlightCurrentEls(testDiv, i);
    for (let j = 0; j < 5; j++) {
        let b = getComputedStyle(children[j]).border;
        if (j == i || j == i + 1) {            
            assert.match(b, redBorderRegex);
            assert.isAbove(parseFloat(b), 0);
        } else {
            assert.notMatch(b, redBorderRegex);
        }
    }
    revertBorder();
}
testDiv.remove();
```

Quando você clicar em `#generate-btn`, deve preencher `#starting-array` com cinco elementos `span`, cada um com um número aleatório entre `1` e `100` como seu texto.

```js
const genBtn = document.querySelector("#generate-btn");
genBtn.click();
const children = document.querySelector("#starting-array").querySelectorAll("span");
assert.lengthOf(children, 5);
Array.from(children).forEach(el => {
    assert.equal(el.tagName, "SPAN");
    const num = Number(el.innerText.trim());
    assert.isAtMost(num, 100);
    assert.isAtLeast(num, 1);
})
```

Quando `#starting-array` já contém um array gerado, ou `#array-container` contém o array ordenado, clicar no `#generate-btn` deve remover os outros elementos em `#array-container`, deixando apenas `#starting-array` com números recém-gerados.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
genBtn.dispatchEvent(new Event("click"));

const prevNumbers = Array.from(document.querySelector("#starting-array").querySelectorAll("span")).map(el => Number(el.innerText.trim()));

sortBtn.dispatchEvent(new Event("click"));
genBtn.dispatchEvent(new Event("click"));

const container = document.querySelector("#array-container");
assert.lengthOf(container.children, 1);
const numbers = Array.from(document.querySelector("#starting-array").querySelectorAll("span")).map(el => Number(el.innerText.trim()));
assert.lengthOf(numbers, 5);
assert.isTrue(prevNumbers.some((num, index) => num !== numbers[index]))
```

Após clicar em `#sort-btn`, `#array-container` deve conter tantos elementos `div` quantos os passos necessários pelo algoritmo Bubble Sort para ordenar o array inicial, incluindo o `div` que representa o array inicial e um `div` que representa o array ordenado.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
// using randomMocker to be sure that the starting array requires 13 steps to be sorted
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    const container = document.querySelector("#array-container");
    assert.lengthOf(container.children, 13)
    Array.from(container.children).forEach(el => {assert.equal(el.tagName, "DIV")})
} finally {
    randomMocker.restore();
}
```

Após clicar em `#sort-btn`, cada `div` dentro de `#array-container` deve conter cinco `span`, cada um com um número como seu texto, e organizados para representar os passos necessários pelo algoritmo Bubble Sort para ordenar o array inicial.

```js
const finalArr = [9, 23, 26, 38, 58];
const arrays = [
    [26, 9, 58, 23, 38],
    [9, 26, 58, 23, 38],
    [9, 26, 58, 23, 38],
    [9, 26, 23, 58, 38],
    [9, 26, 23, 38, 58],
    [9, 26, 23, 38, 58],
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr
]
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    const container = document.querySelector("#array-container");
    assert.isNotEmpty(container.children);
    Array.from(container.children).forEach((el, i) => {
        Array.from(el.children).forEach((j, k) => {
            assert.strictEqual(Number(j.innerText.trim()), arrays[i][k])
        })
    })
} finally {
    randomMocker.restore();
}
```

Quando você clicar no `#sort-btn`, deve utilizar a função `highlightCurrentEls` para destacar os elementos que estão sendo comparados em cada etapa.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
let flag = false;
const temp = highlightCurrentEls;
highlightCurrentEls = () => flag = true;
try {
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    assert.isTrue(flag);
} finally {
    highlightCurrentEls = temp;
}
```

Após clicar em `#sort-btn`, `#starting-array` deve representar o passo inicial com o array inicial e os dois primeiros inteiros destacados usando `highlightCurrentEls`.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");   
genBtn.dispatchEvent(new Event("click"));
sortBtn.dispatchEvent(new Event("click"));
const firstContainer = document.querySelector("#starting-array");
const children = firstContainer.children
const redBorderRegex = /dashed (rgb\(255,\s*0,\s*0\)|#FF0000|hsl\(0,\s*100%,\s*50%\))/;
assert.match(getComputedStyle(children[0]).border, redBorderRegex);
assert.match(getComputedStyle(children[1]).border, redBorderRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <meta http-equiv="X-UA-Compatible" content="ie=edge">
    <title>Sorting Visualizer</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <div id="array-container">
            <div id="starting-array"></div>
        </div>
        <div id="btn-container">
            <button id="generate-btn" type="button">Generate Array</button>
            <button id="sort-btn" type="button">Sort Array</button>
        </div>
    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

main {
    height: 100vh;
    display: flex;
    justify-content: center;
    flex-direction: column;
    align-items: center;
}

#array-container {
    max-height: 95vh;
    display: flex;
    flex-direction: column;
    flex-wrap: wrap;
    gap: 2px;

}

#array-container>div {
    min-width: 8rem;
    height: 2rem;
    box-shadow: rgba(50, 50, 93, 0.25) 0px 2px 5px -1px, rgba(0, 0, 0, 0.3) 0px 1px 3px -1px;
    border-radius: 10px;
    margin-bottom: 0.2rem;
    border: 2px solid darkgray;
    display: flex;
    justify-content: space-evenly;
    align-items: center;
}

#starting-array {
    border: 4px solid darkblue !important;
}

#btn-container {
    display: flex;
    justify-content: space-around;
}

button {
    padding: 2px;
    margin: 5px;
}

span {
    border-radius: 2px;
    padding: 0.5px;
    margin: 0
}

@media (min-width: 430px) {
  #array-container>div {
    min-width: 12rem;    
  }
  span {
    padding: 1px;
    margin: 1px;
  }
}
```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <meta http-equiv="X-UA-Compatible" content="ie=edge">
    <title>Sorting Visualizer</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <div id="array-container">
            <div id="starting-array"></div>
        </div>
        <div id="btn-container">
            <button id="generate-btn" type="button">Generate Array</button>
            <button id="sort-btn" type="button">Sort Array</button>
        </div>
    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

main {
    height: 100vh;
    display: flex;
    justify-content: center;
    flex-direction: column;
    align-items: center;
}

#array-container {
    max-height: 95vh;
    display: flex;
    flex-direction: column;
    flex-wrap: wrap;
    gap: 2px;

}

#array-container>div {
    min-width: 8rem;
    height: 2rem;
    box-shadow: rgba(50, 50, 93, 0.25) 0px 2px 5px -1px, rgba(0, 0, 0, 0.3) 0px 1px 3px -1px;
    border-radius: 10px;
    margin-bottom: 0.2rem;
    border: 2px solid darkgray;
    display: flex;
    justify-content: space-evenly;
    align-items: center;
}

#starting-array {
    border: 4px solid darkblue !important;
}

#btn-container {
    display: flex;
    justify-content: space-around;
}

button {
    padding: 2px;
    margin: 5px;
}

#sort-btn {
   display: none 
}

span {
    border-radius: 2px;
    padding: 0.5px;
    margin: 0
}

@media (min-width: 430px) {
  #array-container>div {
    min-width: 12rem;    
  }
  span {
    padding: 1px;
    margin: 1px;
  }
}
```

```js
const arrayContainer = document.getElementById("array-container");
const generateArrayBtn = document.getElementById("generate-btn");
const sortArrayBtn = document.getElementById("sort-btn");
const arrayLength = 5;
const minVal = 1;
const maxVal = 100;
const isStart = () => arrayContainer.children.length === 1;
const clearArrayContainer = () => { arrayContainer.innerHTML = '<div id="starting-array"></div>' };
const showSortBtn = () => {sortArrayBtn.style.display = "inline-block"};
const hideSortBtn = () => {sortArrayBtn.style.display = "none"};
const generateElement = () => Math.floor(Math.random() * maxVal + minVal);
const generateArray = () => Array.from({ length: arrayLength }, generateElement);
const fillArrContainer = (container, arr) => {
    container.innerHTML = "";
    arr.forEach(i => {
        const el = document.createElement("span");
        el.innerText = i;
        el.id = `number-${i}`
        container.appendChild(el);
    })
}
const isOrdered = (el1, el2) => el1 <= el2;

const swapElements = (arr, n = 0) => {
    if (n < arr.length - 1 && !isOrdered(arr[n], arr[n + 1])) {
        const temp = arr[n];
        arr[n] = arr[n + 1];
        arr[n + 1] = temp;
    }
}
const getLastChildren = () => arrayContainer.lastElementChild;
const getLastArr = () => {
    const els = Array.from(getLastChildren().children);
    const arr = els.map(el => Number(el.id.replace("number-", "")));
    return arr;
}
const generateContainer = () => {
    const container = document.createElement("div");
    arrayContainer.appendChild(container);
    return container;
}
const highlightCurrentEls = (container, n = 0) => {
    const children = container.children;
    children[n].style.border = "2px dashed red";
    children[n + 1].style.border = "2px dashed red";
}

const highlightSorted = () => {
    getLastChildren().style.border = "4px solid green";
}

const bubbleSort = () => {
    let swapped = true;
    while (swapped) {
        const startingArr = getLastArr();
        startingArr.forEach((_, i) => {
            if (i + 1 < arrayLength) {
                highlightCurrentEls(getLastChildren(), i);
                const arr = getLastArr();
                swapElements(arr, i);
                fillArrContainer(generateContainer(), arr);
            }
        })
        const lastArr = getLastArr()
        if (startingArr.every((el, i) => el === lastArr[i])) swapped = false;
    }
}


generateArrayBtn.addEventListener("click", () => {
    if (!isStart()) {
        clearArrayContainer();
    }
    fillArrContainer(document.getElementById("starting-array"), generateArray());
    showSortBtn();
});

sortArrayBtn.addEventListener("click", () => {
    bubbleSort();
    highlightSorted();
    hideSortBtn();
})
```

