---
id: 587d824f367417b2b2512c5a
title: Chai-HTTP を使って API レスポンスの機能テストを実行する III - PUT メソッド
challengeType: 2
forumTopicId: 301590
dashedName: run-functional-tests-on-an-api-response-using-chai-http-iii---put-method
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> からクローンした以下のスタータープロジェクトを基に構築されています。

`PUT` リクエストをテストするときは、多くの場合データを一緒に送信します。`PUT` リクエストに含めるデータはリクエストのボディと呼ばれます。

`PUT` リクエストと JSON オブジェクトを `'/travellers'` エンドポイントに送信するには、`chai-http` プラグインの `put` と `send` メソッドを使うことができます。

```js
chai
  .request(server)
  .keepOpen()
  .put('/travellers')
  .send({
    "surname": [last name of a traveller of the past]
  })
  ...
```

そしてルートは以下のようにレスポンスを返します。

```json
{
  "name": [first name],
  "surname": [last name],
  "dates": [birth - death years]
}
```

`'/travellers'` エンドポイントに対する異なるレスポンスはサーバーコードを参照してください。

# --instructions--

`tests/2_functional-tests.js` 内で、`'Send {surname: "Colombo"}'` テスト（`// #3`）を変更し、`put` と `send` メソッドを使って `'/travellers'` エンドポイントをテストしてください。

PUT リクエストに以下の JSON オブジェクトを送信してください。

```json
{
  "surname": "Colombo"
}
```

`request.end` コールバック内で以下を確認してください。

1. `status` は `200` であること
2. `type` は `application/json` であること
3. `body.name` は `Cristoforo` であること
4. `body.surname` は `Colombo` であること

上記のアサーションの順序を守ってください。これに依存しています。また、完了したら `assert.fail()` を必ず削除してください。

# --hints--

すべてのテストが通過する必要があります。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

`res.status` が 200 であることをテストしてください。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

`res.type` が `'application/json'` であることをテストしてください。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.type');
assert.match(data.assertions[1].args[1], /('|")application\/json\1/);
```

`res.body.name` が `'Cristoforo'` であることをテストしてください。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'equal');
assert.equal(data.assertions[2].args[0], 'res.body.name');
assert.match(data.assertions[2].args[1], /('|")Cristoforo\1/);
```

`res.body.surname` が `'Colombo'` であることをテストしてください。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'equal');
assert.equal(data.assertions[3].args[0], 'res.body.surname');
assert.match(data.assertions[3].args[1], /('|")Colombo\1/);
```
