---
id: 587d824c367417b2b2512c4c
title: ".deepEqual と .notDeepEqual で厳密な等値性をアサートする"
challengeType: 2
forumTopicId: 301587
dashedName: assert-deep-equality-with--deepequal-and--notdeepequal
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> からクローンした以下のスタータープロジェクトを基に作成されています。

`deepEqual()` は2つのオブジェクトが厳密な等値性を持つことをアサートします。

# --instructions--

`tests/1_unit-tests.js` スイート内の `#7` とラベル付けされたテストの `Equality` 内で、各 `assert` を `assert.deepEqual` または `assert.notDeepEqual` のいずれかに変更してテストを通過させてください（`true` と評価される必要があります）。アサートに渡される引数は変更しないでください。

# --hints--

すべてのテストが通過する必要があります。

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
})();
```

最初のアサーションには正しいメソッドを選択してください - `deepEqual` と `notDeepEqual` のどちらかです。

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'deepEqual',
    'The order of the keys does not matter'
  );
})();
```

2番目のアサーションには正しいメソッドを選択してください - `deepEqual` と `notDeepEqual` のどちらかです。

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'notDeepEqual',
    'The position of elements within an array does matter'
  );
})();
```
