---
id: 587d824d367417b2b2512c53
title: 文字列に部分文字列が含まれているかテストする
challengeType: 2
forumTopicId: 301597
dashedName: test-if-a-string-contains-a-substring
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>からクローンした以下のスタータープロジェクトを基に作成されています。

`include()` と `notInclude()` は文字列にも使えます！`include()` は実際の文字列が期待される部分文字列を含んでいることをアサートします。

# --instructions--

`tests/1_unit-tests.js` スイート内の `#14` とラベル付けされたテストの `Strings` 内で、各 `assert` を `assert.include` または `assert.notInclude` のいずれかに変更してテストを通過させてください（`true` と評価されるはずです）。アサートに渡される引数は変更しないでください。

# --hints--

すべてのテストが通過するはずです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

最初のアサーションには正しいメソッドを選んでください - `include` と `notInclude` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'include',
  "'Arrow' contains 'row'..."
);
```

2つ目のアサーションには正しいメソッドを選んでください - `include` と `notInclude` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'notInclude',
  "... a 'dart' doesn't contain a 'queue'"
);
```
