---
id: 587d8250367417b2b2512c5d
title: ヘッドレスブラウザーを使った機能テストの実行
challengeType: 2
forumTopicId: 301595
dashedName: run-functional-tests-using-a-headless-browser
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>からクローンした以下のスタータープロジェクトを基に構築されています。

ページには入力フォームがあります。このフォームはAJAXリクエストとして`PUT /travellers`エンドポイントにデータを送信します。

リクエストが正常に完了すると、クライアントコードはレスポンス内の情報を含む`<div>`をDOMに追加します。

以下はZombie.jsを使ってフォームとやり取りする例です：

```js
test('Submit the surname "Polo" in the HTML form', function (done) {
  browser.fill('surname', 'Polo').then(() => {
    browser.pressButton('submit', () => {
      browser.assert.success();
      browser.assert.text('span#name', 'Marco');
      browser.assert.text('span#surname', 'Polo');
      browser.assert.elements('span#dates', 1);
      done();
    });
  });
});
```

まず、`fill`オブジェクトの`browser`メソッドがフォームの`surname`フィールドに`'Polo'`の値を入力します。`fill`はPromiseを返すので、`then`がチェーンされます。

`then`コールバック内で、`pressButton`オブジェクトの`browser`メソッドを使ってフォームの`submit`イベントリスナーを呼び出します。`pressButton`メソッドは非同期です。

その後、AJAXリクエストからレスポンスが返されると、以下のアサーションが行われます：

1.  レスポンスのステータスが`200`であること
2.  `<span id='name'></span>`要素内のテキストが`'Marco'`と一致すること
3.  `<span id='surname'></span>`要素内のテキストが`'Polo'`と一致すること
4.  `1`要素が`<span id='dates'></span>`存在すること

最後に、非同期テストのために必要な`done`コールバックが呼び出されます。

# --instructions--

`tests/2_functional-tests.js`内の`'Submit the surname "Colombo" in the HTML form'`テスト（`// #5`）で、以下を自動化してください：

1.  フォームに姓`Colombo`を入力する
2.  送信ボタンを押す

そして`pressButton`コールバック内で：

1.  ステータスがOKであることをアサートする`200`
2.  `span#name`要素内のテキストが`'Cristoforo'`であることをアサートする
3.  `span#surname`要素内のテキストが`'Colombo'`であることをアサートする
4.  `span#dates`要素が存在し、その数が`1`であることをアサートする

`assert.fail()`の呼び出しを忘れずに削除してください。

# --hints--

すべてのテストが通過するはずです。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

ヘッドレスブラウザーのリクエストが成功したことをアサートしてください。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'browser.success');
```

`span#name`要素内のテキストが`'Cristoforo'`であることをアサートしてください。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'browser.text');
assert.match(data.assertions[1].args[0], /('|")span#name\1/);
assert.match(data.assertions[1].args[1], /('|")Cristoforo\1/);
```

`span#surname`要素内のテキストが`'Colombo'`であることをアサートしてください。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'browser.text');
assert.match(data.assertions[2].args[0], /('|")span#surname\1/);
assert.match(data.assertions[2].args[1], /('|")Colombo\1/);
```

`span#dates`要素が存在し、その数が1であることをアサートしてください。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'browser.elements');
assert.match(data.assertions[3].args[0], /('|")span#dates\1/);
assert.equal(data.assertions[3].args[1], 1);
```
