---
id: 587d824f367417b2b2512c5b
title: Chai-HTTP を使った API レスポンスの機能テスト IV - PUT メソッド
challengeType: 2
forumTopicId: 301591
dashedName: run-functional-tests-on-an-api-response-using-chai-http-iv---put-method
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> からクローンした以下のスタータープロジェクトを基に構築されています。

この演習は前回のものと似ています。

`PUT` リクエストのテスト方法がわかったので、今度は最初から自分でやってみましょう。

# --instructions--

`tests/2_functional-tests.js` 内で、`'Send {surname: "da Verrazzano"}'` の `// #4` テストを変更し、`put` と `send` メソッドを使って `'/travellers'` エンドポイントをテストしてください。

PUT リクエストに以下の JSON オブジェクトを送信してください。

```json
{
  "surname": "da Verrazzano"
}
```

`request.end` コールバック内で以下を確認してください。

1. `status` は `200` であること
2. `type` は `application/json` であること
3. `body.name` は `Giovanni` であること
4. `body.surname` は `da Verrazzano` であること

上記のアサーションの順序を守ってください。これに依存しています。また、完了したら `assert.fail()` を必ず削除してください。

# --hints--

すべてのテストが通るようにしてください

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

`res.status` が 200 であることをテストしてください

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

`res.type` が `'application/json'` であることをテストしてください

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.type');
assert.match(data.assertions[1].args[1], /('|")application\/json\1/);
```

`res.body.name` が `'Giovanni'` であることをテストしてください

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'equal');
assert.equal(data.assertions[2].args[0], 'res.body.name');
assert.match(data.assertions[2].args[1], /('|")Giovanni\1/);
```

`res.body.surname` が `'da Verrazzano'` であることをテストしてください

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'equal');
assert.equal(data.assertions[3].args[0], 'res.body.surname');
assert.match(data.assertions[3].args[1], /('|")da Verrazzano\1/);
```
