---
id: 587d824d367417b2b2512c52
title: 文字列かどうかをテストする
challengeType: 2
forumTopicId: 301599
dashedName: test-if-a-value-is-a-string
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>からクローンした以下のスタータープロジェクトを基に作成されています。

`isString` または `isNotString` は、実際の値が文字列であることをアサートします。

# --instructions--

`tests/1_unit-tests.js` スイート内の `#13` とラベル付けされたテストの `Strings` 内で、各 `assert` を `assert.isString` または `assert.isNotString` のいずれかに変更してテストを通過させてください（`true` と評価される必要があります）。アサートに渡される引数は変更しないでください。

# --hints--

すべてのテストが通過する必要があります。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

最初のアサーションには正しいメソッドを選んでください - `isString` と `isNotString` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isNotString',
  'A float number is not a string'
);
```

2番目のアサーションには正しいメソッドを選んでください - `isString` と `isNotString` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'isString',
  'environment vars are strings (or undefined)'
);
```

3番目のアサーションには正しいメソッドを選んでください - `isString` と `isNotString` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'isString', 'A JSON is a string');
```
