---
id: 587d824a367417b2b2512c46
title: JavaScriptのアサーションの仕組みを学ぼう
challengeType: 2
forumTopicId: 301589
dashedName: learn-how-javascript-assertions-work
---

# --description--

これらのチャレンジに取り組むには、次のいずれかの方法でコードを書いてください。

- <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">このGitHubリポジトリ</a>をクローンして、ローカルでチャレンジを完了する。
- お好きなサイトビルダーを使ってプロジェクトを完成させる。必ずGitHubリポジトリのすべてのファイルを取り込んでください。

# --instructions--

`tests/1_unit-tests.js`スイート内の`#1`とラベル付けされたテストの`Basic Assertions`の中で、各`assert`を`assert.isNull`または`assert.isNotNull`のいずれかに変更してテストを通過させてください（`true`と評価される必要があります）。アサートに渡される引数は変更しないでください。

# --hints--

すべてのテストが通過する必要があります。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

最初のアサーションには正しいメソッドを選んでください - `isNull`か`isNotNull`のどちらかです。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isNull', 'Null is null');
```

2番目のアサーションには正しいメソッドを選んでください - `isNull`か`isNotNull`のどちらかです。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'isNotNull', '1 is not null');
```
