---
id: 587d824c367417b2b2512c4f
title: 特定の範囲内に値があるかテストする
challengeType: 2
forumTopicId: 301598
dashedName: test-if-a-value-falls-within-a-specific-range
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>からクローンした以下のスタータープロジェクトを基に作成されています。

```javascript
.approximately(actual, expected, delta, [message])
```

`actual`が`expected`と等しいことを、±`delta`の範囲内でアサーションします。

# --instructions--

`tests/1_unit-tests.js`スイートの`#10`というラベルのテスト内の`Comparisons`で、各`assert`を`assert.approximately`に変更してテストを通過させてください（`true`と評価されるはずです）。

テストが常に通過するように、最小値の範囲（3番目のパラメータ）を選んでください。1より小さくする必要があります。

# --hints--

すべてのテストが通過するはずです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

最初のアサーション - `approximately(actual, expected, range)` に対して正しい範囲を選んでください。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'approximately');
assert.equal(
  data.assertions[0].args[2],
  0.5,
  "weirdNumbers(0.5) is in the range (0.5, 1.5]. It's within 1 +/- 0.5"
);
```

2番目のアサーション - `approximately(actual, expected, range)` に対して正しい範囲を選んでください。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'approximately');
assert.equal(
  data.assertions[1].args[2],
  0.8,
  "weirdNumbers(0.2) is in the range (0.2, 1.2]. It's within 1 +/- 0.8"
);
```
