---
id: 587d824b367417b2b2512c4b
title: 厳密な等価性をアサートするために三重等号を使う
challengeType: 2
forumTopicId: 301610
dashedName: use-the-triple-equals-to-assert-strict-equality
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>からクローンした以下のスタータープロジェクトを基に構築されています。

`strictEqual()` は `===` を使ってオブジェクトを比較します。

# --instructions--

`tests/1_unit-tests.js` スイートの中の `#6` とラベル付けされたテスト内の `Equality` において、各 `assert` を `assert.strictEqual` または `assert.notStrictEqual` に変更してテストを通過させてください（`true` と評価される必要があります）。アサートに渡される引数は変更しないでください。

# --hints--

すべてのテストが通過する必要があります。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

最初のアサーションには正しいメソッドを選んでください - `strictEqual` と `notStrictEqual` のどちらかです。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'notStrictEqual',
    'with strictEqual the type must match'
  );
```

2番目のアサーションには正しいメソッドを選んでください - `strictEqual` と `notStrictEqual` のどちらかです。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'strictEqual', '3*2 = 6...');
```

3番目のアサーションには正しいメソッドを選んでください - `strictEqual` と `notStrictEqual` のどちらかです。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'strictEqual',
    "6 * '2' is 12. Types match !"
  );
```

4番目のアサーションには正しいメソッドを選んでください - `strictEqual` と `notStrictEqual` のどちらかです。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[3].method,
    'notStrictEqual',
    'Even if they have the same elements, the Arrays are notStrictEqual'
  );
```
