---
id: 587d824b367417b2b2512c47
title: 変数または関数が定義されているかテストする
challengeType: 2
forumTopicId: 301602
dashedName: test-if-a-variable-or-function-is-defined
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>からクローンした以下のスタータープロジェクトを基に作成されています。

# --instructions--

`tests/1_unit-tests.js`スイート内の`#2`とラベル付けされたテストの`Basic Assertions`内で、各`assert`を`assert.isDefined()`または`assert.isUndefined()`のいずれかに変更してテストを通過させてください（`true`と評価される必要があります）。アサーションに渡される引数は変更しないでください。

# --hints--

すべてのテストが通過する必要があります。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

最初のアサーションには正しいメソッドを選んでください - `isDefined` と `isUndefined` のどちらかです。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'isDefined',
    'Null is not undefined'
  );
```

2番目のアサーションには正しいメソッドを選んでください - `isDefined` と `isUndefined` のどちらかです。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'isUndefined',
    'Undefined is undefined'
  );
```

3番目のアサーションには正しいメソッドを選んでください - `isDefined` と `isUndefined` のどちらかです。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'isDefined',
    'A string is not undefined'
  );
```
