---
id: 587d824b367417b2b2512c49
title: 真偽値のテスト
challengeType: 2
forumTopicId: 301596
dashedName: test-for-truthiness
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>からクローンした以下のスタータープロジェクトを基に作成されています。

`isTrue()` はブール値 `true` をテストし、`isNotTrue()` はブール値 `true` 以外の値が与えられた場合に成功します。

```js
assert.isTrue(true, 'This will pass with the boolean value true');
assert.isTrue('true', 'This will NOT pass with the string value "true"');
assert.isTrue(1, 'This will NOT pass with the number value 1');
```

`isFalse()` と `isNotFalse()` も存在し、それぞれ対応する true のものと似た動作をしますが、ブール値 `false` を探します。

# --instructions--

`tests/1_unit-tests.js` 内の `#4` スイートにある `Basic Assertions` とラベル付けされたテストで、各 `assert` を `assert.isTrue` または `assert.isNotTrue` に変更してテストを通過させてください（`true` と評価されるべきです）。assert に渡される引数は変更しないでください。

# --hints--

すべてのテストが通過する必要があります。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

最初のアサーションには正しいメソッドを選んでください - `isTrue` と `isNotTrue` のどちらかです。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isTrue', 'True is true');
```

2番目のアサーションには正しいメソッドを選んでください - `isTrue` と `isNotTrue` のどちらかです。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'isTrue',
    'Double negation of a truthy value is true'
  );
```

3番目のアサーションには正しいメソッドを選んでください - `isTrue` と `isNotTrue` のどちらかです。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'isNotTrue',
    'A truthy object is not true - neither is a false one'
  );
```
