---
id: 587d824e367417b2b2512c55
title: オブジェクトにプロパティがあるかテストする
challengeType: 2
forumTopicId: 301604
dashedName: test-if-an-object-has-a-property
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>からクローンした以下のスタータープロジェクトを基に作成されています。

`property` は、実際のオブジェクトに指定されたプロパティがあることをアサートします。

# --instructions--

`tests/1_unit-tests.js` スイート内の `#16` とラベル付けされたテストの `Objects` 内で、各 `assert` を `assert.property` または `assert.notProperty` のいずれかに変更してテストを通過させてください（`true` と評価される必要があります）。アサートに渡される引数は変更しないでください。

# --hints--

すべてのテストが通過する必要があります。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

最初のアサーションには正しいメソッドを選んでください - `property` と `notProperty` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notProperty',
  'A car has not wings'
);
```

2番目のアサーションには正しいメソッドを選んでください - `property` と `notProperty` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'property',
  'planes have engines'
);
```

3番目のアサーションには正しいメソッドを選んでください - `property` と `notProperty` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'property', 'Cars have wheels');
```
