---
id: 587d824e367417b2b2512c56
title: 特定のデータ構造型かどうかをテストする
challengeType: 2
forumTopicId: 301601
dashedName: test-if-a-value-is-of-a-specific-data-structure-type
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>からクローンした以下のスタータープロジェクトを基に構築されています。

`#typeOf` は、`Object.prototype.toString` によって判定された値の型が指定された文字列であることをアサートします。

# --instructions--

`tests/1_unit-tests.js`スイート内の`#17`とラベル付けされたテストの`Objects`内で、各`assert`を`assert.typeOf`または`assert.notTypeOf`のいずれかに変更してテストを通過させてください（`true`と評価されるはずです）。アサートに渡される引数は変更しないでください。

# --hints--

すべてのテストが通過する必要があります。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

最初のアサーションには正しいメソッドを選んでください - `typeOf` と `notTypeOf` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'typeOf',
  'myCar is typeOf Object'
);
```

2番目のアサーションには正しいメソッドを選んでください - `typeOf` と `notTypeOf` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'typeOf',
  'Car.model is a String'
);
```

3番目のアサーションには正しいメソッドを選んでください - `typeOf` と `notTypeOf` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[2].method,
  'notTypeOf',
  'Plane.wings is a Number (not a String)'
);
```

4番目のアサーションには正しいメソッドを選んでください - `typeOf` と `notTypeOf` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'typeOf',
  'Plane.engines is an Array'
);
```

5番目のアサーションには正しいメソッドを選んでください - `typeOf` と `notTypeOf` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[4].method,
  'typeOf',
  'Car.wheels is a Number'
);
```
