---
id: 587d824b367417b2b2512c48
title: Assert.isOK と Assert.isNotOK の使い方
challengeType: 2
forumTopicId: 301607
dashedName: use-assert-isok-and-assert-isnotok
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> からクローンした以下のスタータープロジェクトを基に作成されています。

`isOk()` は真値をテストし、`isNotOk()` は偽値をテストします。

真値と偽値についてもっと学びたい場合は、<a href="https://www.freecodecamp.org/learn/javascript-algorithms-and-data-structures/basic-algorithm-scripting/falsy-bouncer" target="_blank" rel="noopener noreferrer nofollow">Falsy Bouncer</a> チャレンジを試してください。

# --instructions--

`tests/1_unit-tests.js` スイート内の `#3` とラベル付けされたテストの `Basic Assertions` の中で、各 `assert` を `assert.isOk()` または `assert.isNotOk()` に変更してテストを通過させてください（`true` と評価される必要があります）。assert に渡される引数は変更しないでください。

# --hints--

すべてのテストが通過する必要があります。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

最初のアサーションには正しいメソッドを選んでください - `isOk` と `isNotOk` のどちらかです。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isNotOk', 'Null is falsy');
```

2番目のアサーションには正しいメソッドを選んでください - `isOk` と `isNotOk` のどちらかです。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'isOk', 'A string is truthy');
```

3番目のアサーションには正しいメソッドを選んでください - `isOk` と `isNotOk` のどちらかです。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[2].method, 'isOk', 'true is truthy');
```
