---
id: 587d824d367417b2b2512c51
title: 配列にアイテムが含まれているかテストする
challengeType: 2
forumTopicId: 301603
dashedName: test-if-an-array-contains-an-item
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>からクローンした以下のスタータープロジェクトを基に作成されています。

# --instructions--

`tests/1_unit-tests.js`スイート内の`#12`とラベル付けされたテストの`Arrays`の中で、各`assert`を`assert.include`か`assert.notInclude`のどちらかに変更してテストを通過させてください（`true`と評価されるはずです）。assertに渡されている引数は変更しないでください。

# --hints--

すべてのテストが通過する必要があります。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

最初のアサーションには正しいメソッドを選んでください - `include` と `notInclude` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notInclude',
  "It's summer in july..."
);
```

2つ目のアサーションには正しいメソッドを選んでください - `include` と `notInclude` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'include',
  'JavaScript is a back-end language !!'
);
```
