---
id: 587d824d367417b2b2512c53
title: 文字列に部分文字列が含まれているかどうかをテストする
challengeType: 2
forumTopicId: 301597
dashedName: test-if-a-string-contains-a-substring
---

# --description--

As a reminder, this project is being built upon the following starter project cloned from <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`include()` と `notInclude()` は、文字列でも動作します！ `include()` は、実際の文字列が期待される部分文字列を含んでいることをアサートします。

# --instructions--

Within `tests/1_unit-tests.js` under the test labeled `#14` in the `Strings` suite, change each `assert` to either `assert.include` or `assert.notInclude` to make the test pass (should evaluate to `true`). アサートに渡された引数を変更しないでください。

# --hints--

すべてのテストが成功する必要があります。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

最初のアサーションに対して、正しいメソッドを選ぶ必要があります - `include` もしくは `notInclude` です。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'include',
  "'Arrow' contains 'row'..."
);
```

2 番目のアサーションに対して、正しいメソッドを選ぶ必要があります - `include` もしくは `notInclude` です。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'notInclude',
  "... a 'dart' doesn't contain a 'queue'"
);
```

