---
id: 587d824e367417b2b2512c57
title: コンストラクタのインスタンスかどうかをテストする
challengeType: 2
forumTopicId: 301605
dashedName: test-if-an-object-is-an-instance-of-a-constructor
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>からクローンした以下のスタータープロジェクトを基に作成されています。

`#instanceOf` はオブジェクトがコンストラクタのインスタンスであることをアサートします。

# --instructions--

`tests/1_unit-tests.js`スイート内の`#18`というラベルの付いたテストの`Objects`内で、各`assert`を`assert.instanceOf`か`assert.notInstanceOf`のいずれかに変更してテストを通過させてください（`true`と評価されるはずです）。アサートに渡される引数は変更しないでください。

# --hints--

すべてのテストが通過する必要があります。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

最初のアサーションには正しいメソッドを選んでください - `instanceOf` と `notInstanceOf` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notInstanceOf',
  'myCar is not an instance of Plane'
);
```

2番目のアサーションには正しいメソッドを選んでください - `instanceOf` と `notInstanceOf` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'instanceOf',
  'airlinePlane is an instance of Plane'
);
```

3番目のアサーションには正しいメソッドを選んでください - `instanceOf` と `notInstanceOf` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[2].method,
  'instanceOf',
  'everything is an Object in JavaScript...'
);
```

4番目のアサーションには正しいメソッドを選んでください - `instanceOf` と `notInstanceOf` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'notInstanceOf',
  'myCar.wheels is not an instance of String'
);
```
