---
id: 587d824d367417b2b2512c50
title: 値が配列かどうかをテストする
challengeType: 2
forumTopicId: 301600
dashedName: test-if-a-value-is-an-array
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>からクローンした以下のスタータープロジェクトを基に構築されています。

# --instructions--

`tests/1_unit-tests.js`スイート内の`#11`とラベル付けされたテストの`Arrays`の中で、各`assert`を`assert.isArray`または`assert.isNotArray`のいずれかに変更してテストを通過させてください（`true`と評価される必要があります）。アサートに渡される引数は変更しないでください。

# --hints--

すべてのテストが通過する必要があります。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=10');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

最初のアサーションには正しいメソッドを選んでください - `isArray` と `isNotArray` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=10');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isArray',
  'String.prototype.split() returns an Array'
);
```

2つ目のアサーションには正しいメソッドを選んでください - `isArray` と `isNotArray` のどちらかです。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=10');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'isNotArray',
  'Array.prototype.indexOf() returns a number'
);
```
