---
id: 587d824f367417b2b2512c59
title: Chai-HTTP II を使って API エンドポイントの機能テストを実行する
challengeType: 2
forumTopicId: 301592
dashedName: run-functional-tests-on-api-endpoints-using-chai-http-ii
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> からクローンした以下のスタータープロジェクトを基に構築されています。

# --instructions--

`tests/2_functional-tests.js` 内で、テスト `'Test GET /hello with your name'`（`// #2`）を変更して、レスポンスの `status` と `text` をアサートし、テストを通過させてください。

名前は URL クエリとして、ルートに `?name=<your_name>` を追加して送信してください。エンドポイントは `'hello <your_name>'` でレスポンスします。

# --hints--

すべてのテストが通過する必要があります

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=1');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

`res.status` が 200 であることをテストしてください

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=1');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

`res.text` が `'hello <your_name>'` と等しいことをテストしてください

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=1');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.text');
assert.match(data.assertions[1].args[1], /hello [\w\d_-]/);
```
