---
id: 5f8884f4c46685731aabfc41
title: ヘッドレスブラウザーを使った機能テストの実行 II
challengeType: 2
forumTopicId: 301594
dashedName: run-functional-tests-using-a-headless-browser-ii
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>からクローンした以下のスタータープロジェクトを基に構築されています。

# --instructions--

`tests/2_functional-tests.js`内の`'Submit the surname "Vespucci" in the HTML form'`テスト（`// #6`）で、次の操作を自動化してください。

1. フォームに姓`Vespucci`を入力する
2. 送信ボタンを押す

そして`pressButton`コールバック内で：

1. ステータスがOKであることをアサートする `200`
2. 要素`span#name`内のテキストが`'Amerigo'`であることをアサートする
3. 要素`span#surname`内のテキストが`'Vespucci'`であることをアサートする
4. 要素`span#dates`が存在し、その数が`1`であることをアサートする

`assert.fail()`の呼び出しを削除するのを忘れないでください。

# --hints--

すべてのテストが通過する必要があります。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

ヘッドレスブラウザーのリクエストが成功したことをアサートしてください。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'browser.success');
```

要素`span#name`内のテキストが`'Amerigo'`であることをアサートしてください。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'browser.text');
assert.match(data.assertions[1].args[0], /('|")span#name\1/);
assert.match(data.assertions[1].args[1], /('|")Amerigo\1/);
```

要素`span#surname`内のテキストが`'Vespucci'`であることをアサートしてください。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'browser.text');
assert.match(data.assertions[2].args[0], /('|")span#surname\1/);
assert.match(data.assertions[2].args[1], /('|")Vespucci\1/);
```

要素`span#dates`が存在し、その数が1であることをアサートしてください。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'browser.elements');
assert.match(data.assertions[3].args[0], /('|")span#dates\1/);
assert.equal(data.assertions[3].args[1], 1);
```
