---
id: 66db529d37ad966480ebb633
title: Створіть програму вибору обіду
challengeType: 26
dashedName: build-a-lunch-picker-program
---

# --description--

У цій лабораторній роботі ви створите програму, яка допомагає керувати варіантами обіду. Ви працюватимете з масивом обідів, додаватимете та видалятимете елементи з масиву, а також випадково вибиратимете варіант обіду.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні оголосити змінну `lunches` і призначити їй порожній масив, який буде використовуватися для зберігання елементів обіду.

2. Ви повинні створити функцію `addLunchToEnd`, яка приймає масив як перший аргумент і рядок як другий аргумент. Функція повинна:
    - Додати рядок у кінець масиву.
    - Вивести в консоль рядок `[Lunch Item] added to the end of the lunch menu.`, де `[Lunch Item]` — це рядок, переданий у функцію.
    - Повернути оновлений масив.

3. Ви повинні створити функцію `addLunchToStart`, яка приймає масив як перший аргумент і рядок як другий аргумент. Функція повинна:
    - Додати рядок на початок масиву.
    - Вивести в консоль рядок `[Lunch Item] added to the start of the lunch menu.`, де `[Lunch Item]` — це рядок, переданий у функцію.
    - Повернути оновлений масив.

4. Ви повинні створити функцію `removeLastLunch`, яка приймає масив як аргумент. Функція повинна:
   - Видалити останній елемент із масиву.
   - Якщо видалення відбулося успішно, вивести в консоль рядок `[Lunch Item] removed from the end of the lunch menu.`, де `[Lunch Item]` — це елемент, видалений із масиву.
   - Якщо масив порожній, вивести в консоль рядок `"No lunches to remove."`.
   - Повернути оновлений масив.

5. Ви повинні створити функцію `removeFirstLunch`, яка приймає масив як аргумент. Функція повинна:
   - Видалити перший елемент із масиву.
   - Якщо видалення відбулося успішно, вивести в консоль рядок `[Lunch Item] removed from the start of the lunch menu.`, де `[Lunch Item]` — це елемент, видалений із масиву.
   - Якщо масив порожній, вивести в консоль рядок `"No lunches to remove."`.
   - Повернути оновлений масив.

6. Ви повинні створити функцію `getRandomLunch`, яка приймає масив як аргумент. Функція повинна:
   - Вибрати випадковий елемент із масиву.
   - Якщо вибір відбувся успішно, вивести в консоль рядок `Randomly selected lunch: [Lunch Item]`, де `[Lunch Item]` — випадковий елемент із масиву.
   - Якщо масив порожній, вивести в консоль рядок `"No lunches available."`.

7. Ви повинні створити функцію `showLunchMenu`, яка приймає масив як аргумент і:
   - Якщо в масиві є елементи, виводить у консоль рядок `Menu items: [Lunch Item], [Lunch Item]...`, де кожен `[Lunch item]` — це один із елементів масиву, у порядку.
   - Якщо масив порожній, виводить у консоль рядок `"The menu is empty."`.

# --hints--

Ви повинні оголосити змінну `lunches` і призначити їй порожній масив для зберігання елементів обіду.

```js
assert.isArray(lunches);
```

Ви повинні визначити функцію `addLunchToEnd`.

```js
assert.isFunction(addLunchToEnd);
```

Функція `addLunchToEnd` повинна мати два параметри.

```js
assert.lengthOf(addLunchToEnd, 2);
```

`addLunchToEnd(lunches, "Tacos")` має виводити рядок `"Tacos added to the end of the lunch menu."` у консоль.

```js
const tempArr = [];
const temp = console.log;
const testLunches = []
try {
  console.log = obj => tempArr.push(obj);
  addLunchToEnd(testLunches, "Tacos");
  assert.strictEqual(tempArr[0], "Tacos added to the end of the lunch menu.");
  addLunchToEnd(testLunches, "Pizza");
  assert.strictEqual(tempArr[1], "Pizza added to the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToEnd(["Pizza", "Tacos"], "Burger")` має повертати `["Pizza", "Tacos", "Burger"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToEnd(["Pizza", "Tacos"], "Burger"), ["Pizza", "Tacos", "Burger"]);
  assert.deepStrictEqual(addLunchToEnd(["Fries"], "Tacos"), ["Fries", "Tacos"]);
} finally {
  console.log = temp;
}
```

Ви повинні визначити функцію `addLunchToStart`.

```js
assert.isFunction(addLunchToStart);
```

Функція `addLunchToStart` повинна мати два параметри.

```js
assert.lengthOf(addLunchToStart, 2);
```

`addLunchToStart(lunches, "Sushi")` має виводити рядок `"Sushi added to the start of the lunch menu."` у консоль.

```js
const tempArr = [];
const testLunches = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  addLunchToStart(testLunches, "Sushi");
  assert.strictEqual(tempArr[0], "Sushi added to the start of the lunch menu.");  
  addLunchToStart(testLunches, "Burger");
  assert.strictEqual(tempArr[1], "Burger added to the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToStart(["Burger", "Sushi"], "Pizza")` має повертати `["Pizza", "Burger", "Sushi"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToStart(["Burger", "Sushi"], "Pizza"), ["Pizza", "Burger", "Sushi"]);
  assert.deepStrictEqual(addLunchToStart(["Noodles"], "Sushi"), ["Sushi", "Noodles"]);
} finally {
  console.log = temp;
}
```

Ви повинні визначити функцію `removeLastLunch`.

```js
assert.isFunction(removeLastLunch)
```

Функція `removeLastLunch` повинна мати один параметр.

```js
assert.lengthOf(removeLastLunch, 1);
```

Коли вхідний масив порожній, функція `removeLastLunch` має виводити рядок `"No lunches to remove."` у консоль.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Stew", "Soup", "Toast"])` має виводити рядок `"Toast removed from the end of the lunch menu."` у консоль.

```js
let testLunches = ["Stew", "Soup", "Toast"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "Toast removed from the end of the lunch menu.");
  testLunches = ["Rice", "Pizza"];
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[1], "Pizza removed from the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Sushi", "Pizza", "Noodles"])` має повертати `["Sushi", "Pizza"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeLastLunch(["Sushi", "Pizza", "Noodles"]), ["Sushi", "Pizza"]);
  assert.deepStrictEqual(removeLastLunch(["Pizza", "Burger"]), ["Pizza"]);
} finally {
  console.log = temp;
}
```

Ви повинні визначити функцію `removeFirstLunch`.

```js
assert.isFunction(removeFirstLunch);
```

Функція `removeFirstLunch` повинна мати один параметр.

```js
assert.lengthOf(removeFirstLunch, 1);
```

Коли вхідний масив порожній, функція `removeFirstLunch` має виводити рядок `"No lunches to remove."` у консоль.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Salad", "Eggs", "Cheese"])` має виводити рядок `"Salad removed from the start of the lunch menu."` у консоль.

```js
let testLunches = ["Salad", "Eggs", "Cheese"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "Salad removed from the start of the lunch menu.");
  testLunches = ["Stew", "Soup", "Toast"]
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[1], "Stew removed from the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Sushi", "Pizza", "Burger"])` має повертати `["Pizza", "Burger"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeFirstLunch(["Sushi", "Pizza", "Burger"]), ["Pizza", "Burger"]);
  assert.deepStrictEqual(removeFirstLunch(["Pizza", "Burger"]), ["Burger"]);
} finally {
  console.log = temp;
}
```

`addLunchToEnd`, `addLunchToStart`, `removeLastLunch` і `removeFirstLunch` мають повертати той самий масив, переданий як аргумент, після оновлення.

```js
const temp = console.log;
console.log = () => {}
const testLunches = ["Spam"];
try {
  assert.equal(addLunchToEnd(testLunches, "Pizza"), testLunches);
  assert.equal(addLunchToStart(testLunches, "Caviar"), testLunches);
  assert.equal(removeLastLunch(testLunches), testLunches);
  assert.equal(removeFirstLunch(testLunches), testLunches);
} finally {
  console.log = temp;
}
```

Ви повинні визначити функцію `getRandomLunch`.

```js
assert.isFunction(getRandomLunch);
```

Функція `getRandomLunch` повинна мати один параметр.

```js
assert.lengthOf(getRandomLunch, 1);
```

Коли вхідний масив порожній, функція `getRandomLunch` має виводити рядок `"No lunches available."` у консоль.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  getRandomLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches available.");
} finally {
  console.log = temp;
}
```

Коли вхідний масив не порожній, функція `getRandomLunch` має виводити рядок у форматі `Randomly selected lunch: [Lunch Item]` у консоль.

```js
const testLunches = ["Fish", "Fries", "Roast"];
const tempRandom = Math.random;
const tempArr = [];
const temp = console.log;

try {
  console.log = obj => tempArr.push(obj);
  // check that it correctly outputs the first item
  Math.random = () => 0;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[0], `Randomly selected lunch: ${testLunches[0]}`);

  // Second item
  Math.random = () => 0.4;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[1], `Randomly selected lunch: ${testLunches[1]}`);

  // Third item
  Math.random = () => 0.8;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[2], `Randomly selected lunch: ${testLunches[2]}`);

} finally {
  Math.random = tempRandom;
  console.log = temp;
}
```

Функція `getRandomLunch` не повинна змінювати масив, переданий їй як аргумент.

```js
const testLunches = ["Fish", "Fries", "Roast"];
let originalLunches = [...testLunches];
const tempRandom = Math.random;

try {
  Math.random = () => 0;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.4;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.8;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  const emptyLunches = [];
  let originalEmpty = [...emptyLunches];
  getRandomLunch(emptyLunches);
  assert.sameOrderedMembers(emptyLunches, originalEmpty);

} finally {
  Math.random = tempRandom;
}
```

Ви повинні визначити функцію `showLunchMenu`.

```js
assert.isFunction(showLunchMenu);
```

Функція `showLunchMenu` повинна мати один параметр.

```js
assert.lengthOf(showLunchMenu, 1);
```

Коли вхідний масив порожній, функція `showLunchMenu` має виводити рядок `"The menu is empty."` у консоль.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], "The menu is empty.");
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Greens", "Corns", "Beans"])` має виводити `"Menu items: Greens, Corns, Beans"` у консоль.

```js
const testLunches = ["Greens", "Corns", "Beans"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Pizza", "Burger", "Fries", "Salad"])` має виводити `"Menu items: Pizza, Burger, Fries, Salad"` у консоль.

```js
const testLunches = ["Pizza", "Burger", "Fries", "Salad"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const lunches = [];

// Add a new lunch to the end of the list
function addLunchToEnd(lunches, newLunch) {
  lunches.push(newLunch);
  console.log(`${newLunch} added to the end of the lunch menu.`);
  return lunches;
}

// Add a new lunch to the start of the list
function addLunchToStart(lunches, newLunch) {
  lunches.unshift(newLunch);
  console.log(`${newLunch} added to the start of the lunch menu.`);
  return lunches;
}

// Remove the last lunch from the list
function removeLastLunch(lunches) {
  const removedLunch = lunches.pop();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the end of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Remove the first lunch from the list
function removeFirstLunch(lunches) {
  const removedLunch = lunches.shift();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the start of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Function to get a random lunch from the list
function getRandomLunch(lunches) {
  if (lunches.length === 0) {
    console.log("No lunches available.");
  }
  else {
    const randomIndex = Math.floor(Math.random() * lunches.length);
    const randomLunch = lunches[randomIndex];
    console.log(`Randomly selected lunch: ${randomLunch}`);
  }
}

// Function to display all the lunches in the list
function showLunchMenu(lunches) {
  if (lunches.length === 0) {
    console.log("The menu is empty.");
  }
  else {
    console.log(`Menu items: ${lunches.join(', ')}`);
  }
}

addLunchToEnd(lunches, "Tacos");
addLunchToStart(lunches, "Sushi");
removeLastLunch(lunches);
removeFirstLunch(lunches);
getRandomLunch(lunches);
showLunchMenu(lunches);

```
