---
id: 65578d0f6c78a0b868a43b9c
title: 步驟 40
challengeType: 20
dashedName: step-40
---

# --description--

在 `while` 循環內，首先要做的是定義要訪問的當前節點。 爲此你可以使用 `min()` 函數。 它從作爲參數傳遞的可迭代對象中返回最小的項。

移除 `pass` ，然後創建一個名爲 `current` 的變量，並將其賦值爲 `min(unvisited)`。

# --hints--

You should create a `current` variable in your `while` loop.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].has_variable("current")`)) })
```

You should assign `min(unvisited)` to your `current` variable. Remember to delete `pass`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].find_bodies()[0].is_equivalent("current = min(unvisited)")`)) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--
    while unvisited:
        pass
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')

#shortest_path(my_graph, 'A')

```
