---
id: 655777060d8ddea6741be1b1
title: 步驟 32
challengeType: 20
dashedName: step-32
---

# --description--

所有在 `distances` 中的距離都被設置爲無窮大，除了起始節點。 `unvisited` 列表包含了你的圖中的所有節點。 但實際上，您不需要那個`for`循環來達到這個結果。

移除你的 `for` 循環的整個主體。

# --hints--

你應當移除你的 `for` 循環和所有嵌套其中的代碼。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/(^\s*)distances\s*=\s*\{\s*\}\s*\1print\s*\(\s*f("|')Unvisited:\s*\{\s*unvisited\s*\}\\nDistances:\s\{\s*distances\s*\}\2\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
        else:
            distances[node] = float('inf')
    print(f'Unvisited: {unvisited}\nDistances: {distances}')

shortest_path(my_graph, 'A')
--fcc-editable-region--
```
