---
id: 65b7d3319ebec69b983fb91b
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

As you can see from the output, `.items()` creates a data structure that stores each key-value pair in a distinct tuple. 要迭代這些元組中的元素，你可以添加第二個循環變量：

```py
for i, j in dict.items():
    print(i, j)
```

修改你的 `for` 循環以採用兩個循環變量並將它們兩個都打印在循環體內。

# --hints--

你應該修改 `for` 循環以使用兩個變量。

```js
({ test: () => assert(runPython(`
import ast
var = _Node(_code).find_for_loops()[0].find_for_vars()
isinstance(var.tree, ast.Tuple)
`)) })
```

你應該在 `for` 循環內打印兩個循環變量。

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print{var}")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.items():
    print(i)
--fcc-editable-region--
```
