---
id: 65b7cd2b7bd9a684ccf1dc16
title: 步驟 9
challengeType: 20
dashedName: step-9
---

# --description--

要遍歷字典的鍵，你只需將字典放入 `for` 循環中。 下面的代碼將打印字典 `dict` 中的每個鍵：

```py
for i in dict:
   print(i)
```

將 `print()` 調用替換爲 `for` 循環，該循環遍歷 `copper` 並打印每個鍵。

# --hints--

你的代碼中不應該有 `print(copper)`。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(copper)")`)) })
```

你應該創建一個 `for` 循環來遍歷 `copper` 字典。

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper")`)) })
```

你應該在 `for` 循環內打印 `copper` 的每個鍵。

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print({var})")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

print(copper)
--fcc-editable-region--
```
