---
id: 5d822fd413a79914d39e98cf
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

你可以看到 `body` （它是你的页面最里面的​​框）；它周围的框是 `html` 元素。 通过给它一个 `100vh` 的 `height` 来让你的 `body` 填满整个视口。 通过将 `margin` 设置为 `0`，从 `body` 中删除默认的 `margin`。 最后，将 `overflow` 属性设置为 `hidden` 以隐藏任何超出视口时出现的滚动条。

# --hints--

应该使用 `body` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

你的 `body` 应该有一个 `100vh` 的 `height`。

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

你的 `body` 应该有一个 `0` 的 `margin`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

你的 `body` 应该有设置为 `hidden` 的 `overflow`属性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--
    
```

