---
id: 5d822fd413a79914d39e98cc
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

欢迎来到 CSS 变量天际线项目！ 首先在 `head` 元素中添加一个 `link` 元素链接你的 `styles.css` 文件。

# --hints--

你的代码中应该有一个 `link` 元素。

```js
assert.isNotNull(document.querySelector('link'));
```

你不应更改现有的 `head` 标签。 确保你没有删除结束标签。

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

你应该有一个 `link` 空元素。

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

`link` 元素应该在 `head` 元素中。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

你的 `link` 元素应具有一个 `rel` 属性，其值为 `stylesheet`。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

你的 `link` 元素应具有一个 `href` 属性，其值为 `styles.css`。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
