---
id: 5d822fd413a79914d39e98d4
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

给你的建筑的各个部分添加 `width` 和 `height` 属性：`.bb1a` 的宽度为 `70%`，高度为 `10%`；`.bb1b` 的宽度为 `80%`，高度为 `10%`；`.bb1c` 的宽度为 `90%`，高度为 `10%`；`.bb1d` 的宽度为 `100%`，高度为 `70%`。 请记住，这些百分比是相对于父级的，并注意高度将加起来为 100%——垂直填充容器。

# --hints--

你应该使用类选择器来设置 `.bb1a` 的样式。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a'));
```

你应该给 `.bb1a` 添加一个 `70%` 的 `width`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.width, '70%');
```

你应该给 `.bb1a` 添加一个 `10%` 的 `height`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.height, '10%');
```

你应该使用类选择器来设置 `.bb1b` 的样式。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b'));
```

你应该给 `.bb1b` 添加一个 `80%` 的 `width`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.width, '80%');
```

你应该给 `.bb1b` 添加一个 `10%` 的 `height`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.height, '10%');
```

你应该使用类选择器来设置 `.bb1c` 的样式。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c'));
```

你应该给 `.bb1c` 添加一个 `90%` 的 `width`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.width, '90%');
```

你应该给 `.bb1c` 添加一个 `10%` 的 `height`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.height, '10%');
```

你应该使用类选择器来设置 `.bb1d` 的样式。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d'));
```

你应该给 `.bb1d` 添加一个 `100%` 的 `width`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.width, '100%');
```

你应该给 `.bb1d` 添加一个 `70%` 的 `height`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
--fcc-editable-region--

--fcc-editable-region--

```

