---
id: 6909ed80e1c195f49564edac
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Neste workshop, você aprenderá como trabalhar com declarações condicionais e operadores de comparação.

Comece declarando uma variável chamada `truthyOrFalsy` e atribua a ela o valor booleano `true`.

# --hints--

Você deve ter uma variável chamada `truthyOrFalsy`.

```js
assert.exists(truthyOrFalsy);
```

Você deve atribuir `true` à sua variável `truthyOrFalsy`.

```js
assert.isTrue(truthyOrFalsy);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
