---
id: 690a57095c4afb27bbe751b3
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Valores falsy são valores que avaliam para `false` em um contexto booleano. Exemplos de valores falsy são:

- ""
- `0`
- `false`
- `null`
- `undefined`
- `NaN`

Você encontrará todos estes ao longo deste curso.

Por enquanto, altere a atribuição de `truthyOrFalsy` definindo seu valor como uma string vazia. Depois disso, você verá `false` registrado no console.

# --hints--

Você deve atribuir uma string vazia a `truthyOrFalsy`.

```js
assert.strictEqual(truthyOrFalsy, "");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "freeCodeCamp";
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
