---
id: 690a0df27ee9ee36d935d530
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

Declare outra variável chamada `isTimmyAGamer` e atribua a ela o booleano `false`.

# --hints--

Você deve declarar uma variável chamada `isTimmyAGamer`.

```js
assert.exists(isTimmyAGamer);
```

Você deve atribuir o valor booleano `false` à sua variável `isTimmyAGamer`.

```js
assert.isFalse(isTimmyAGamer);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}
--fcc-editable-region--

--fcc-editable-region--
```
