---
id: 690a58cc4f828a49eab1b53a
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Declare uma variável chamada `hasDeveloperJob` e atribua a ela o valor booleano `true`.

# --hints--

Você deve ter uma variável chamada `hasDeveloperJob`.

```js
assert.exists(hasDeveloperJob);
```

Você deve atribuir `true` à sua variável `hasDeveloperJob`.

```js
assert.isTrue(hasDeveloperJob);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
