---
id: 690a06833d3e83b5226ade78
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Uma declaração condicional é usada para executar código dependendo de uma condição. Ela é composta pela palavra-chave `if`, seguida por uma condição entre parênteses e um corpo entre chaves. Aqui está um exemplo de uma declaração `if`:

```js
if (condition) {
  console.log("condition is truthy");
}
```

Quando `condition` é um valor truthy (string não vazia, número diferente de zero, `true`, etc.) o código dentro do corpo da declaração `if` é executado.

Abaixo da sua variável, crie uma declaração `if` que use `hasDeveloperJob` como sua condição. Dentro do corpo da sua declaração `if`, registre `"Timmy is employed as a developer."` no console.

# --hints--

Você deve ter uma instrução `if`.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(.+\)\s*\{?/);
```

Sua declaração `if` deve usar `hasDeveloperJob` como sua condição.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{?/);
```

Você deve registrar `"Timmy is employed as a developer."` no console dentro do seu `if` statement.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is employed as a developer.\1\s*\)\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const hasDeveloperJob = true;


--fcc-editable-region--
```
