---
id: 690a51806fb743bc896c3af7
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Como você aprendeu em uma lição anterior, booleanos são um tipo de dado com apenas os valores `true` e `false`.

Um valor truthy é um valor que avalia para `true` em um contexto booleano. Exemplos de valores truthy são:

- Strings não vazias
- Números diferentes de zero
- O booleano `true`

E outros valores que você encontrará mais adiante.

Você pode usar a função `Boolean()` para verificar a veracidade de um valor. Por exemplo, `Boolean("Hello World!")` retornará `true` porque `"Hello World!"` é truthy.

Agora, chame `Boolean(truthyOrFalsy)` e registre o resultado no console.

# --hints--

Você deve registrar `Boolean(truthyOrFalsy)` no console.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const truthyOrFalsy = true;
--fcc-editable-region--

--fcc-editable-region--
```
