---
id: 690a57fef66aac39dffc61a2
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Agora que você sabe o que são valores truthy e falsy, remova tanto a declaração da sua variável quanto a instrução `console.log` do seu código.

# --hints--

Você não deve ter uma variável `truthyOrFalsy`.

```js
assert.notMatch(__helpers.removeJSComments(code), /const\s+truthyOrFalsy\s*=/);
```

Você não deve ter `console.log(Boolean(truthyOrFalsy))` no seu código.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "";

console.log(Boolean(truthyOrFalsy));
--fcc-editable-region--
```
