---
id: 690a3ce03013be44d116136e
title: Passo 12
challengeType: 1
dashedName: step-12
---

# --description--

Uma declaração condicional pode ter uma cláusula `else`, que executa código quando a condição `if` é falsa. Aqui está um exemplo de uma declaração `if...else`:

```js
if (condition) {
  console.log("condition is truthy");
} else {
  console.log("condition is falsy");
}
```

Adicione uma cláusula `else` à declaração `if` existente. Dentro do corpo da sua cláusula `else`, registre `"Timmy is not old enough to drive."` no console.

# --hints--

Sua terceira declaração `if` deve ter uma cláusula `else`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}\s+else\s+{/s);
```

Você deve registrar `"Timmy is not old enough to drive."` no console dentro do corpo da sua cláusula `else`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}\s+else\s+{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is not old enough to drive.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}
--fcc-editable-region--
const timmyAge = 18;

if (timmyAge >= 16) {
  console.log("Timmy is old enough to drive.");
}
--fcc-editable-region--
```
