---
id: 690a0f7a5c22c15222a6aa5c
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

Agora crie outra declaração `if` que use `isTimmyAGamer` como sua condição. Dentro do corpo da sua declaração `if`, registre `"Timmy loves to play World of Warcraft."` no console.

Observe que você não verá nada registrado no console porque a condição do seu novo `if` statement` é falsa.

# --hints--

Você deve ter duas declarações `if` no seu código.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(.+\)\s*\{.+\}/s);
```

Sua segunda declaração `if` deve usar `isTimmyAGamer` como sua condição.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{.+\}/s);
```

Você deve registrar `"Timmy loves to play World of Warcraft."` no console dentro do seu segundo `if` statement.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy loves to play World of Warcraft.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

--fcc-editable-region--

--fcc-editable-region--
```
