---
id: 690a562fd611e8155599ce98
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Altere a atribuição de `truthyOrFalsy` definindo seu valor para `"freeCodeCamp"`. Você ainda verá `true` registrado no console porque uma string não vazia é um valor truthy.

# --hints--

Sua variável `truthyOrFalsy` deve ter o valor `"freeCodeCamp"`.

```js
assert.strictEqual(truthyOrFalsy, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = true;
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
