---
id: 690a3ad44874a928a2fe3077
title: Passo 11
challengeType: 1
dashedName: step-11
---

# --description--

Operadores de comparação permitem que você compare valores e escreva declarações condicionais mais complexas.

```js
if (today == birthday) {
  console.log("Happy Birthday!");
}
```

Por exemplo, o código acima registra `"Happy Birthday!"` no console somente quando os valores das variáveis `today` e `birthday` são iguais.

Crie uma terceira declaração `if`. Para sua condição, use o operador `>=` para verificar se `timmyAge` é maior ou igual a `16`.

Dentro do corpo do seu novo `if` statement, registre `"Timmy is old enough to drive."` no console.

# --hints--

Você deve ter uma terceira declaração `if` no seu código.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(.+\)\s*\{.+\}/s);
```

Sua terceira declaração `if` deve usar `timmyAge >= 16` como sua condição.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s* 16\s*\)\s*\{.+\}/s);
```

Você deve registrar `"Timmy is old enough to drive."` no console dentro do seu terceiro `if` statement.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}
--fcc-editable-region--
const timmyAge = 18;


--fcc-editable-region--
```
