---
id: 690a3f884e083a6a33b81fcf
title: Passo 13
challengeType: 1
dashedName: step-13
---

# --description--

Finalmente, altere a declaração de `timmyAge` e defina-a para um número menor que `16`. Depois disso você verá `"Timmy is not old enough to drive."` registrado no console.

# --hints--

Sua variável `timmyAge` deve ter um valor menor que `16`.

```js
assert.isAtMost(timmyAge, 15);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}
--fcc-editable-region--
const timmyAge = 18;
--fcc-editable-region--
if (timmyAge >= 16) {
  console.log("Timmy is old enough to drive.");
} else {
  console.log("Timmy is not old enough to drive.");
}
```

# --solutions--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 15;

if (timmyAge >= 16) {
  console.log("Timmy is old enough to drive.");
} else {
  console.log("Timmy is not old enough to drive.");
}
```
