---
id: 690a110a76d0496d9e7c28eb
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

Declare outra variável chamada `timmyAge` e defina-a como `18`.

# --hints--

Você deve ter uma variável chamada `timmyAge`.

```js
assert.exists(timmyAge);
```

Você deve atribuir `18` à sua variável `timmyAge`.

```js
assert.strictEqual(timmyAge, 18);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}
--fcc-editable-region--

--fcc-editable-region--
```
