---
id: 66faac4139dbbd5dd632c860
title: بناء نموذج طلب وظيفة
challengeType: 25
dashedName: lab-job-application-form
demoType: onClick
---

# --description--

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك عنصر `div` بالصنف `container`.
1. داخل عنصر `div`، يجب أن يكون لديك عنصر `form`.
1. يجب أن يحتوي النموذج على عنصر `input` بالنوع `text` والمعرّف `name` لإدخال الاسم الكامل للمستخدم.
1. يجب أن يكون لديك عنصر `input` آخر بالنوع `email` والمعرّف `email` لإدخال عنوان البريد الإلكتروني للمستخدم.
1. يجب أن يتضمن النموذج عنصر `select` بالمعرّف `position` يتيح للمستخدمين اختيار منصب وظيفي.
1. يجب أن يكون لديك عنصر `fieldset` بصنف `radio-group`.
1. داخل `.radio-group`، يجب أن يكون لديك مجموعة من عناصر `input` بالنوع `radio` مع تسميات ملائمة لاختيار خيارات التوفر (مثل: دوام كامل، دوام جزئي). وينبغي أن يكون `name` المجموعة `availability`.
1. يجب أن يكون لديك عنصر `textarea` بالمعرّف `message` لإدخال رسالة.
1. يجب أن تربط كل عنصر `input` و`select` و`textarea` بعنصر `label`.
1. يجب أن يكون لديك عنصر `button` بالنوع `submit` لإرسال النموذج.
1. أضف الفئة الزائفة `:focus` إلى عنصري `input` و`textarea` لتغيير لون حدودهما وإزالة المخطط الافتراضي عند التركيز.
1. يجب أن تحمل عناصر `input` و`select` و`textarea` فئة زائفة `:invalid` تغيّر لون الحدود إلى الأحمر عند اكتشاف إدخال غير صالح.
1. يجب أن تحمل عناصر `input` و`select` و`textarea` فئة زائفة `:valid` تغيّر لون الحدود إلى الأخضر عند إدخال مدخل صالح.
1. يجب أن يحمل عنصر `button` فئة زائفة `:hover` تغيّر لون الخلفية عند المرور بالمؤشر فوقه.
1. استخدم الفئة الزائفة `:checked` على `.radio-group input[type="radio"]` لإضافة لون حدود ولون خلفية وظل صندوق عند تحديد زر الاختيار.
1. استخدم الفئة الزائفة `:checked` على أزرار الاختيار لتغيير لون نص `label` المرتبط عند تحديد الخيار.
1. أضف الفئة الزائفة `:first-of-type` إلى عنصر `input` لتنسيق أول حقل إدخال بشكل مختلف. (مثل: زوايا مستديرة).

**ملاحظة:** تأكد من ربط ورقة الأنماط في ملف HTML وتطبيق CSS الخاص بك.

# --hints--

يجب أن يكون لديك عنصر `div` بالصنف `container`.

```js
assert.exists(document.querySelector("div.container"));
```

داخل عنصر `div`، يجب أن يكون لديك عنصر `form`.

```js
assert.exists(document.querySelector("div.container > form"));
```

يجب أن يحتوي النموذج على عنصر `input` من النوع `text` وبالمعرف `name` لإدخال الاسم الكامل للمستخدم.

```js
assert.exists(document.querySelector("div.container > form input#name[type='text']"));
```

يجب أن يكون لديك عنصر `input` آخر من النوع `email` وبالمعرف `email` لإدخال عنوان البريد الإلكتروني للمستخدم.

```js
assert.exists(document.querySelector("div.container > form input#email[type='email']"));
```

يجب أن يتضمن النموذج عنصر `select` بالمعرف `position` مع بعض عناصر `option`.

```js
assert.exists(document.querySelector("div.container > form select#position"));
assert.isNotEmpty(document.querySelectorAll("div.container > form select#position > option"))
```

يجب أن يكون لديك عنصر `fieldset` أو `section` بالصنف `.radio-group`.

```js
assert.lengthOf(document.querySelectorAll('form fieldset.radio-group, form section.radio-group'), 1)
```

داخل `.radio-group` يجب أن يحتوي على مجموعة من عناصر `input` من النوع `radio` لاختيار خيارات التوفر. يجب أن يكون اسم المجموعة `name` هو `availability`.

```js
assert.isAtLeast(document.querySelectorAll("div.container > form .radio-group input[type='radio'][name='availability']").length, 2);
assert.isEmpty(document.querySelectorAll("div.container > form .radio-group input[type='radio']:not([name='availability'])"))
```

يجب أن يكون لديك عنصر `textarea` بالمعرف `message` لإدخال رسالة.

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form textarea#message"));
```

يجب ربط كل عنصر من عناصر `input` و`select` و`textarea` بعنصر `label`.

```js
const els = document.querySelectorAll('input, select, textarea');
assert.isNotEmpty(els);
els.forEach(el => {
  const label = document.querySelector(`label[for="${el.id}"]`) || el.parentElement;
  assert.exists(label); 
  assert.equal(label.tagName, "LABEL");
})
```

يجب أن يكون لديك عنصر `button` من النوع `submit` لإرسال النموذج.

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form button[type='submit']"));
```

يجب إضافة فئة مزيفة `:focus` لعناصر `input` و`textarea` لتغيير لون الحدود عند التركيز. استخدم محدد قائمة بالترتيب المعطى.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input:focus, textarea:focus').getPropertyValue('border-color'));
```

يجب استخدام محدد فئة مزيفة `:focus` لعناصر `input` و`textarea` أيضًا لإزالة الإطار الافتراضي عند التركيز.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:focus, textarea:focus').getPropertyValue('outline'), 'none');
```

يجب أن تحتوي عناصر `input` و`select` و`textarea` على فئة مزيفة `:invalid` تغير لون الحدود إلى الأحمر عند اكتشاف إدخال غير صالح. استخدم محدد قائمة بالترتيب المعطى.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:invalid, select:invalid, textarea:invalid').getPropertyValue('border-color'), 'red')
```

يجب أن تحتوي عناصر `input` و`select` و`textarea` على فئة مزيفة `:valid` تغير لون الحدود إلى الأخضر عند إدخال بيانات صحيحة. استخدم محدد قائمة بالترتيب المعطى.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:valid, select:valid, textarea:valid').getPropertyValue('border-color'), 'green')
```

يجب أن يحتوي عنصر `button` على فئة مزيفة `:hover` تغير لون الخلفية عند المرور فوقه.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button:hover').getPropertyValue('background-color'))
```

يجب استخدام فئة مزيفة `:checked` على `.radio-group input[type="radio"]` لإضافة لون حدود عند اختيار زر الاختيار.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('border-color'));
```

يجب استخدام فئة مزيفة `:checked` على `.radio-group input[type="radio"]` لإضافة لون خلفية عند اختيار زر الاختيار.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('background-color'));
```

يجب استخدام فئة مزيفة `:checked` على `.radio-group input[type="radio"]` لإضافة ظل صندوق عند اختيار زر الاختيار.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('box-shadow'));
```

يجب استخدام الفئة المزيفة `:checked` على أزرار الاختيار لتغيير لون نص الـ `label` المرتبط عندما يتم اختيار الخيار.

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.equal(label.tagName, "LABEL");
  input.checked = false;
  const colorBefore = getComputedStyle(label).color;
  input.checked = true;
  const colorAfter = getComputedStyle(label).color;
  assert.notEqual(colorAfter, colorBefore);
})
```

أضف فئة مزيفة `:first-of-type` لعناصر `input` لتنسيق أول حقل إدخال بشكل مختلف.

```js
const pseudo = new __helpers.CSSHelp(document).getStyle('input:first-of-type');
assert.exists(pseudo);
assert.isNotEmpty([...pseudo]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application Form</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="container">
        <h1>Job Application Form</h1>
        <form id="job-application">
            <label for="name">Full Name:</label>
            <input type="text" id="name" name="name" placeholder="Enter your name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" placeholder="Enter your email" required>

            <label for="position">Position:</label>
            <select id="position" name="position" required>
                <option value="">Select a position</option>
                <option value="developer">Developer</option>
                <option value="designer">Designer</option>
                <option value="manager">Manager</option>
            </select>

            <fieldset class="radio-group">
              <legend>Availability:</legend>
                <input type="radio" id="fullTime" name="availability" value="fullTime" checked>
                <label for="fullTime">Full-Time</label>
                <input type="radio" id="partTime" name="availability" value="partTime">
                <label for="partTime">Part-Time</label>
            </fieldset>

            <label for="message">Why do you want this job?</label>
            <textarea id="message" name="message" rows="5" placeholder="Write your motivation" required></textarea>

            <button type="submit">Submit Application</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: white;
  padding: 20px;
}

.container {
  max-width: 600px;
  margin: 0 auto;
  background-color: white;
  padding: 20px;
  border-radius: 8px;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
}

label {
  display: block;
  margin-bottom: 5px;
  font-weight: bold;
}

input, select, textarea {
  width: 100%;
  padding: 10px;
  margin-bottom: 15px;
  border: 1px solid #ccc;
  border-radius: 4px;
  box-sizing: border-box;
}

.radio-group {
  display: flex;
  align-items: center;
  margin-top: 15px;
}
.radio-group input[type="radio"] {
  margin-right: 10px;
  -webkit-appearance: none;
  -moz-appearance: none;
  appearance: none;
  width: 20px;
  height: 20px;
  border: 2px solid #ccc;
  border-radius: 50%;
  outline: none;
  transition: border-color 0.2s;
}

.radio-group label {
  margin: 0;
  font-weight: normal;
  cursor: pointer;
}

button {
  width: 100%;
  padding: 10px;
  background-color: green;
  border: none;
  color: white;
  border-radius: 4px;
  font-size: 16px;
  cursor: pointer;
}

input:focus, textarea:focus {
  border-color: blue;
  outline: none;
}

input:invalid, select:invalid, textarea:invalid {
  border-color: red;
}

input:valid, select:valid, textarea:valid {
  border-color: green;
}

.radio-group input[type="radio"]:checked {
  border-color: green;
  background-color: green;
  box-shadow: inset 0 0 0 4px white;
}

button:hover {
  background-color: black;
}

input[type="radio"]:checked + label {
  color: green;
}

button:disabled {
  background-color: grey;
}

input:first-of-type {
  border-top-left-radius: 4px;
  border-top-right-radius: 4px;
}

textarea:nth-child(5) {
  border-bottom-left-radius: 4px;
  border-bottom-right-radius: 4px;
}
```
