---
id: 612e6afc009b450a437940a1
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Inizia con la struttura HTML di base. Aggiungi una dichiarazione `DOCTYPE` e gli elementi `html`, `head`, `body` e `title`.

Imposta la lingua di questa pagina su inglese. Imposta il `title` su `Piano`.

# --hints--

Il tuo codice dovrebbe contenere il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Dovresti includere uno spazio dopo il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Dovresti definire il tipo di documento come `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Dovresti chiudere la dichiarazione `DOCTYPE` con un `>` dopo il tipo.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Il tuo elemento `html` dovrebbe avere un tag di apertura con un attributo `lang` di `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Il tuo elemento `html` dovrebbe avere un tag di chiusura.

```js
assert(code.match(/<\/html\s*>/));
```

La dichiarazione `DOCTYPE` dovrebbe essere all'inizio del tuo HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Dovresti avere un tag di apertura `head`.

```js
assert(code.match(/<head\s*>/i));
```

Dovresti avere un tag di chiusura `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Dovresti avere un tag di apertura `body`.

```js
assert(code.match(/<body\s*>/i));
```

Dovresti avere un tag di chiusura `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Gli elementi `head` e `body` dovrebbero essere fratelli.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

L'elemento `head` dovrebbe essere all'interno dell'elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

L'elemento `body` dovrebbe essere all'interno dell'elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Il tuo codice dovrebbe avere un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Il tuo progetto dovrebbe avere un titolo `Piano`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'piano')
```

Ricorda, la maiuscola e l'ortografia sono importanti per il titolo.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Piano');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
