---
id: 612e77aba7ca691f598feb02
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Aggiungi due tag `meta`: uno per ottimizzare la tua pagina per dispositivi mobili e uno per specificare un `charset` accettato per la pagina.

# --hints--

Dovresti aggiungere due elementi `meta` alla tua pagina.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Dovresti avere due elementi `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Uno degli elementi `meta` dovrebbe avere un insieme `name` impostato su `viewport` e `content` impostato su `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

L'altro elemento `meta` dovrebbe avere l'attributo `charset` impostato su `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Piano</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body></body>
</html>
```

```css

```
