---
id: 612e78af05201622d4bab8aa
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

È ora di iniziare a lavorare sul piano. Crea un elemento `div` all'interno del tuo elemento `body` con l'insieme `id` impostato su `piano`.

# --hints--

Dovresti creare un nuovo elemento `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Il tuo `div` dovrebbe essere all'interno del tuo elemento `body`.

```js
const div = document.querySelector('div');
assert(div?.parentElement?.localName === 'body');
```

Il tuo `div` dovrebbe avere l'insieme `id` impostato su `piano`.

```js
const div = document.querySelector('div');
assert(div?.getAttribute('id') === 'piano');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css

```
