---
id: 66a9c006e4f58819396605b3
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

head セクションが完成したので、次は <dfn>table body</dfn>、`tbody` を追加しましょう。table body はすべての学生の名前と成績を表します。

テーブルに table body 要素を追加してください。

# --hints--

開く `tbody` タグがあるはずです。

```js
assert.match(code, /<tbody>/i);
```

閉じる `tbody` タグがあるはずです。

```js
assert.match(code, /<\/tbody>/i);
```

開く `tbody` タグは閉じる `table` タグの前に来るはずです。

```js
assert.match(code, /<tbody>[\s\S]*<\/tbody>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
