---
id: 66a9ca099e3ecb1be10a2696
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

これでテーブルにさらに2人の生徒を追加する時間です。

前のステップと同じパターンに従って、4人目の生徒のテーブル行を追加してください。テーブルデータ要素には次のデータを使います：

- 苗字: `Thompson`
- 名前: `Jane`
- 成績: `77`

5人目の生徒のテーブル行には、次のデータを使います：

- 苗字: `Williams`
- 名前: `Natalie`
- 成績: `83`

# --hints--

4番目の`tr`が`tbody`要素の中にあるはずです。

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(4)'));
```

4番目の`td`要素の中に、テキストが`Thompson`の`tr`要素があるはずです。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:first-child')?.textContent.trim(), 'Thompson');
```

4番目の`td`要素の中に、テキストが`Jane`の`tr`要素があるはずです。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(2)')?.textContent.trim(), 'Jane');
```

4番目の`td`要素の中に、テキストが`77`の`tr`要素があるはずです。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(3)')?.textContent.trim(), '77');
```

5番目の`tr`が`tbody`要素の中にあるはずです。

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 5);
```

5番目の`td`要素の中に、テキストが`Williams`の`tr`要素があるはずです。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:first-child')?.textContent.trim(), 'Williams');
```

5番目の`td`要素の中に、テキストが`Natalie`の`tr`要素があるはずです。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(2)')?.textContent.trim(), 'Natalie');
```

5番目の`td`要素の中に、テキストが`83`の`tr`要素があるはずです。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(3)')?.textContent.trim(), '83');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
