---
id: 66a9bcf00f13a418368a272e
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

テーブルのヘッダー要素は、<dfn>テーブル行</dfn>要素である `tr` で構成されており、その中に <dfn>テーブルヘッダーセル</dfn>要素である `th` が含まれています。

スポーツのテーブルで `tr` と `th` 要素を使った例を示します。

```html
<table>
  <caption>Football Scores</caption>
  <thead>
    <tr>
      <th>Team</th>
      <th>Wins</th>
      <th>Losses</th>
    </tr>
  </thead>
</table>
```

`thead` 要素の中に `tr` 要素を追加してください。

`tr` 要素の中に `th` 要素を3つ追加してください。

最初の `th` 要素にはテキスト `Last Name` を入れてください。2番目の `th` 要素にはテキスト `First Name` を入れてください。3番目の `th` 要素にはテキスト `Grade` を入れてください。

# --hints--

開く `tr` タグがあるはずです。

```js
assert.match(code, /<tr>/i);
```

閉じる `tr` タグがあるはずです。

```js
assert.match(code, /<\/tr>/i);
```

`tr` 要素は `thead` 要素の中にあるはずです。

```js
const trEl = document.querySelector('tr');
assert.strictEqual(trEl?.parentElement?.tagName, 'THEAD');
```

最初の `th` 要素には `Last Name` のテキストがあるはずです。

```js
const thEl = document.querySelectorAll('th')[0];
assert.strictEqual(thEl?.textContent.trim(), 'Last Name');
```

2番目の `th` 要素には `First Name` のテキストがあるはずです。

```js
const thEl = document.querySelectorAll('th')[1];
assert.strictEqual(thEl?.textContent.trim(), 'First Name');
```

3番目の `th` 要素には `Grade` のテキストがあるはずです。

```js
const thEl = document.querySelectorAll('th')[2];
assert.strictEqual(thEl?.textContent.trim(), 'Grade');
```

すべての `th` 要素は `tr` 要素の中にあるはずです。

```js
assert.match(code, /<tr>.*<th>.*<\/th>.*<th>.*<\/th>.*<th>.*<\/th>.*<\/tr>/is);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
      --fcc-editable-region--
        
      --fcc-editable-region--
      </thead>
    </table>
  </body>
</html>
```
