---
id: 66a9c14e3b34c719e34bae20
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

テーブルに学生データを追加するには、テーブル行要素とテーブルデータ要素を使う必要があります。

<dfn>table data</dfn> 要素、`td` は、テーブルのセルを作成するために使われます。

スポーツ選手のテーブルで `td` 要素を使う例を示します。

```html
<tr>
  <td>1</td>
  <td>John Doe</td>
  <td>USA</td>
</tr>
```

テーブルの本文要素内に、3つのテーブルデータ要素 `tr` を持つテーブル行要素 `td` を追加してください。

最初のテーブルデータ要素には `Davis` の姓を入れてください。

2番目のテーブルデータ要素には `Alex` の名を入れてください。

3番目のテーブルデータ要素には `54` の成績を入れてください。

# --hints--

`tr` 要素の中に `tbody` 要素があるはずです。

```js
assert.isNotNull(document.querySelector('tbody tr'));
```

`td` 要素の中にテキスト `Davis` を持つ `tr` 要素があるはずです。

```js
assert.strictEqual(document.querySelector('tbody tr td')?.textContent.trim(), 'Davis');
```

`td` 要素の中にテキスト `Alex` を持つ `tr` 要素があるはずです。

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[1]?.textContent.trim(), 'Alex');
```

`td` 要素の中にテキスト `54` を持つ `tr` 要素があるはずです。

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[2]?.textContent.trim(), '54');
```

`td` 要素の中に3つすべての `tr` 要素があるはずです。

```js
assert.match(code, /<tr>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<\/tr>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
