---
id: 66a9cb26ec6bd41cf6c82bc5
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

テーブルに追加する最後のセクションは、<dfn>table foot</dfn> 要素、`tfoot` です。table foot 要素は、テーブル内のすべての学生の平均成績を表示するために使います。

テーブルに `tfoot` 要素を追加してください。

`tfoot` 要素の中に `tr` 要素を追加してください。

`tr` 要素の中に 2つの `td` 要素を追加してください。

最初の `td` 要素にはテキスト `Average Grade` を入れてください。2つ目の `td` 要素にはテキスト `78.8` を入れてください。

# --hints--

開く `tfoot` タグがあるはずです。

```js
assert.match(code, /<tfoot>/);
```

閉じる `tfoot` タグがあるはずです。

```js
assert.match(code, /<\/tfoot>/);
```

`tr` 要素の中に `tfoot` 要素があるはずです。

```js
assert.isNotNull(document.querySelector('tfoot tr'));
```

最初の `td` 要素にはテキスト `Average Grade` が入っているはずです。`tr` 要素の中にあることを確認してください。

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child')?.textContent.trim(), 'Average Grade');
```

2つ目の `td` 要素にはテキスト `78.8` が入っているはずです。`tr` 要素の中にあることを確認してください。

```js
assert.strictEqual(document.querySelector('tfoot tr td:nth-of-type(2)')?.textContent.trim(), '78.8');
```

`tr` 要素の中に 2つの `td` 要素があるはずです。

```js
assert.strictEqual(document.querySelectorAll('tfoot tr td').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
