---
id: 66a9bb8578459a175432e7d0
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

テーブルの最初のセクションでは、生徒のファーストネーム、ラストネーム、期末試験の成績を表す列ラベルのヘッダーコンテンツをグループ化します。

<dfn>table head</dfn> 要素、`thead` は、テーブル内のヘッダーコンテンツをグループ化するために使われます。

以下は `thead` 要素を使った例です。

```html
<table>
  <thead>
    <!-- header content goes here -->
  </thead>
</table>
```

`caption` 要素の下に、table head 要素を追加してください。

# --hints--

開く `thead` タグがあるはずです。

```js
assert.match(code, /<thead>/i);
```

閉じる `thead` タグがあるはずです。

```js
assert.match(code, /<\/thead>/i);
```

`thead` 要素は `table` 要素の内側にあるはずです。

```js
const theadEl = document.querySelector('thead');
assert.strictEqual(theadEl?.parentElement?.tagName, 'TABLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
