---
id: 66a9c7a4521d2b1b1ec6dcf0
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

これでテーブルにさらに2人の生徒を追加する時間です。

前のステップと同じパターンに従って、2人目の生徒のテーブル行を追加してください。テーブルデータ要素には次のデータを使います：

- 苗字: `Doe`
- 名前: `Samantha`
- 成績: `92`

3人目の生徒のテーブル行には、次のデータを使います：

- 苗字: `Rodriguez`
- 名前: `Marcus`
- 成績: `88`

# --hints--

`tr`要素の中に2つ目の`tbody`があるはずです。

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(2)'));
```

2つ目の`td`要素の中にテキスト`Doe`を持つ`tr`要素があるはずです。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:first-child')?.textContent.trim(), 'Doe');
```

2つ目の`td`要素の中にテキスト`Samantha`を持つ`tr`要素があるはずです。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:nth-child(2)')?.textContent.trim(), 'Samantha');
```

2つ目の`td`要素の中にテキスト`92`を持つ`tr`要素があるはずです。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:last-child')?.textContent.trim(), '92');
```

`tr`要素の中に3つ目の`tbody`があるはずです。

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 3);
```

3つ目の`td`要素の中にテキスト`Rodriguez`を持つ`tr`要素があるはずです。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:first-child')?.textContent.trim(), 'Rodriguez');
```

3つ目の`td`要素の中にテキスト`Marcus`を持つ`tr`要素があるはずです。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:nth-child(2)')?.textContent.trim(), 'Marcus');
```

3つ目の`td`要素の中にテキスト`88`を持つ`tr`要素があるはずです。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:last-child')?.textContent.trim(), '88');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
