---
id: 66a9b8f14b963916a3baa732
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`table` にキャプションを追加するには、<dfn>テーブルキャプション</dfn>要素を使うことができます。

`caption` 要素を使った例を示します。

```html
<table>
  <caption>Football Scores</caption>
</table>
```

`table` 要素の中に、テキスト `caption` を持つ `Calculus Final Exam Grades` 要素をネストしてください。

# --hints--

開く `caption` タグがあるはずです。

```js
assert.match(code, /<caption>/i);
```

閉じる `caption` タグがあるはずです。

```js
assert.match(code, /<\/caption>/i);
```

`caption` 要素は `table` 要素の中にあるはずです。

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.parentElement.tagName, 'TABLE');
```

`caption` はテキスト `Calculus Final Exam Grades` を持つはずです。

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.innerText.trim(), 'Calculus Final Exam Grades');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
    --fcc-editable-region--
      
    --fcc-editable-region--      
    </table>
  </body>
</html>
```
