---
id: 66a9b48a211a73155621d0df
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

このワークショップでは、学生グループの期末試験のテーブルを作成しながら、HTMLのテーブル操作を練習します。

これまでのレッスンで、`table` 要素を使って表形式のデータを表す方法を学びました。

`body` 要素の中に、`table` 要素をネストしてください。

# --hints--

開く`table`タグが必要です。

```js
assert.match(code, /<table>/i);
```

閉じる`table`タグが必要です。

```js
assert.match(code, /<\/table>/i);
```

`table` 要素は`body` 要素の中に入っている必要があります。

```js
const tableEl = document.querySelector('table');
assert.strictEqual(tableEl?.parentElement?.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
