---
id: 587d7db2367417b2b2512b89
title: Tumia mixin kuongeza tabia za kawaida kati ya vitu visivyo na uhusiano
challengeType: 1
forumTopicId: 301331
dashedName: use-a-mixin-to-add-common-behavior-between-unrelated-objects
---

# --description--

Kama ulivyoona, tabia hushirikiwa kupitia urithi. Hata hivyo, kuna matukio ambapo urithi si suluhisho bora. Urithi hauendani vizuri kwa vitu visivyo na uhusiano kama `Bird` na `Airplane`. Vyote vinaweza kuruka, lakini `Bird` si aina ya `Airplane` na kinyume chake.

Kwa vitu visivyo na uhusiano, ni bora kutumia <dfn>mixins</dfn>. Mixin huruhusu vitu vingine kutumia mkusanyiko wa kitendakazi.

```js
let flyMixin = function(obj) {
  obj.fly = function() {
    console.log("Flying, wooosh!");
  }
};
```

`flyMixin` huchukua kitu chochote na kinatoa njia ya `fly`.

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let plane = {
  model: "777",
  numPassengers: 524
};

flyMixin(bird);
flyMixin(plane);
```

Hapa `bird` na `plane` hupitishwa ndani ya `flyMixin`, ambayo kisha inawaweka kitendakazi cha `fly` kwa kila kitu. Sasa `bird` na `plane` wote wanaweza kuruka:

```js
bird.fly();
plane.fly();
```

Konsoli itaonyesha mfuatano wa herufi `Flying, wooosh!` mara mbili, mara moja kwa kila wito wa `.fly()`.

Angalia jinsi mixin inavyoruhusu njia ile ile ya `fly` kutumika tena na vitu visivyo na uhusiano `bird` na `plane`.

# --instructions--

Tengeneza mixin iitwayo `glideMixin` inayotangaza njia iitwayo `glide`. Kisha tumia `glideMixin` kumpa `bird` na `boat` uwezo wa kuvinjari.

# --hints--

Msimbo wako unapaswa kutangaza (declare) tofauti ya `glideMixin` ambayo ni kitendakazi.

```js
assert(typeof glideMixin === 'function');
```

Msimbo wako unapaswa kutumia `glideMixin` kwenye kitu cha `bird` kumpa njia ya `glide`.

```js
assert(typeof bird.glide === 'function');
```

Msimbo wako unapaswa kutumia `glideMixin` kwenye kitu cha `boat` kumpa njia ya `glide`.

```js
assert(typeof boat.glide === 'function');
```

# --seed--

## --seed-contents--

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let boat = {
  name: "Warrior",
  type: "race-boat"
};

// Only change code below this line
```

# --solutions--

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let boat = {
  name: "Warrior",
  type: "race-boat"
};
function glideMixin (obj) {
  obj.glide = () => 'Gliding!';
}

glideMixin(bird);
glideMixin(boat);
```
