---
id: 587d7dad367417b2b2512b78
title: Tumia kijenzi kuunda vitu
challengeType: 1
forumTopicId: 18233
dashedName: use-a-constructor-to-create-objects
---

# --description--

Hapa kuna kijenzi cha `Bird` kutoka changamoto iliyopita:

```js
function Bird() {
  this.name = "Albert";
  this.color  = "blue";
  this.numLegs = 2;
}

let blueBird = new Bird();
```

**KUMBUKA:** `this` ndani ya kijenzi daima inarejelea kitu kinachoundwa.

Tambua kuwa opereta ya `new` hutumika wakati wa kuita kijenzi. Hii inamwambia JavaScript kuunda mfano mpya wa `Bird` unaoitwa `blueBird`. Bila opereta ya `new`, `this` ndani ya kijenzi isingerejelea kitu kilichoundwa kipya, na kusababisha matokeo yasiyotegemewa. Sasa `blueBird` ina vigezo vyote vilivyotangazwa ndani ya kijenzi cha `Bird`:

```js
blueBird.name;
blueBird.color;
blueBird.numLegs;
```

Kama kitu kingine chochote, vigezo vyake vinaweza kufikiliwa na kubadilishwa:

```js
blueBird.name = 'Elvira';
blueBird.name;
```

# --instructions--

Tumia kijenzi cha `Dog` kutoka funzo la mwisho kuunda mfano mpya wa `Dog`, ukiuweka kwenye kigezo `hound`.

# --hints--

`hound` inapaswa kuundwa kwa kutumia kijenzi cha `Dog`.

```js
assert(hound instanceof Dog);
```

Msimbo wako unapaswa kutumia opereta ya `new` kuunda mfano wa `Dog`.

```js
assert(__helpers.removeJSComments(code).match(/new/g));
```

# --seed--

## --seed-contents--

```js
function Dog() {
  this.name = "Rupert";
  this.color = "brown";
  this.numLegs = 4;
}
// Only change code below this line
```

# --solutions--

```js
function Dog() {
  this.name = "Rupert";
  this.color = "brown";
  this.numLegs = 4;
}
const hound = new Dog();
```
