---
id: 587d7dad367417b2b2512b77
title: Eleza kitendakazi cha kijenzi
challengeType: 1
forumTopicId: 16804
dashedName: define-a-constructor-function
---

# --description--

<dfn>Kijenzi</dfn> ni vitendakazi vinavyotengeneza vitu vipya. Huvitengenezea sifa na tabia zitakazomilikiwa na kitu kipya. Fikiria kama ramani ya kutengeneza vitu vipya.

Hapa kuna mfano wa kijenzi:

```js
function Bird() {
  this.name = "Albert";
  this.color = "blue";
  this.numLegs = 2;
}
```

Kijenzi hiki huunda kitu cha `Bird` chenye sifa `name`, `color`, na `numLegs` zilizowekwa kuwa Albert, blue, na 2, mtawalia. Kijenzi hufuata mikataba kadhaa:

<ul><li>Vijenzi huandikwa kwa jina lililoanza kwa herufi kubwa ili kutofautisha na vitendakazi vingine ambavyo si <code>vijenzi</code>.</li><li>Vijenzi hutumia neno kuu <code>this</code> kuweka sifa za kitu ambacho kitaunda. Ndani ya kijenzi, <code>this</code> inarejelea kitu kipya kitakachoundwa.</li><li>Vijenzi huainisha sifa na tabia badala ya kurudisha thamani kama vitendakazi vingine vinaweza kufanya.</li></ul>

# --instructions--

Tengeneza kijenzi, `Dog`, chenye sifa `name`, `color`, na `numLegs` ambazo zimewekwa kuwa mfuatano wa herufi, mfuatano wa herufi, na nambari, mtawalia.

# --hints--

`Dog` inapaswa kuwa na sifa ya `name` iliyowekwa kuwa mfuatano wa herufi.

```js
assert(typeof new Dog().name === 'string');
```

`Dog` inapaswa kuwa na sifa ya `color` iliyowekwa kuwa mfuatano wa herufi.

```js
assert(typeof new Dog().color === 'string');
```

`Dog` inapaswa kuwa na sifa ya `numLegs` iliyowekwa kuwa nambari.

```js
assert(typeof new Dog().numLegs === 'number');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function Dog (name, color, numLegs) {
  this.name = 'name';
  this.color = 'color';
  this.numLegs = 4;
}
```
