---
id: 587d7dac367417b2b2512b73
title: Unda kitu cha msingi cha JavaScript
challengeType: 1
forumTopicId: 301317
dashedName: create-a-basic-javascript-object
---

# --description--

Fikiria vitu watu huviona kila siku, kama magari, maduka, na ndege. Hivi ni <dfn>vitu</dfn>: vitu vinavyoonekana na watu wanaweza kuvinigilia na kuingiliana navyo.

Ni sifa gani za vitu hivi? Gari lina magurudumu. Maduka huuza bidhaa. Ndege wana mabawa.

Sifa hizi, au <dfn>sifa</dfn>, zinaelezea kile kinachounda kitu. Kumbuka vitu vinavyofanana vina sifa sawa, lakini vinaweza kuwa na thamani tofauti za sifa hizo. Kwa mfano, magari yote yana magurudumu, lakini si magari yote yana idadi sawa ya magurudumu.

Vitu katika JavaScript hutumika kuiga vitu halisi vya dunia, kwa kuwapa sifa na tabia kama vile vitu halisi. Hapa kuna mfano unaotumia dhana hizi kuunda kitu cha `duck`:

```js
let duck = {
  name: "Aflac",
  numLegs: 2
};
```

Kitu hiki cha `duck` kina jozi mbili za sifa/thamani: `name` ya `Aflac` na `numLegs` ya 2.

# --instructions--

Unda kitu cha `dog` chenye sifa za `name` na `numLegs`, na uzitekee thamani ya mfuatano wa herufi na nambari, mtawaliwa.

# --hints--

`dog` inapaswa kuwa kitu.

```js
assert(typeof dog === 'object');
```

`dog` inapaswa kuwa na sifa ya `name` iliyowekwa kuwa mfuatano wa herufi.

```js
assert(typeof dog.name === 'string');
```

`dog` inapaswa kuwa na sifa ya `numLegs` iliyowekwa kuwa nambari.

```js
assert(typeof dog.numLegs === 'number');
```

# --seed--

## --seed-contents--

```js
let dog = {

};
```

# --solutions--

```js
let dog = {
  name: '',
  numLegs: 4
};
```
