---
id: 587d7db2367417b2b2512b8b
title: Elewa usemi wa kitendakazi kisicho na jina unaoendeshwa mara moja (IIFE)
challengeType: 1
forumTopicId: 301328
dashedName: understand-the-immediately-invoked-function-expression-iife
---

# --description--

Mfumo wa kawaida katika JavaScript ni kuendesha kitendakazi mara tu kinapotangazwa:

```js
(function () {
  console.log("Chirp, chirp!");
})();
```

Huu ni usemi wa kitendakazi kisicho na jina unaoendeshwa mara moja, na hutolewa `Chirp, chirp!` mara moja.

Kumbuka kuwa kitendakazi hakina jina na hakihifadhiwi katika kigezo chochote. Mabano ya kawaida () mawili mwishoni mwa usemi wa kitendakazi husababisha kuendeshwa au kuitwa mara moja. Mfumo huu unajulikana kama <dfn>usemi wa kitendakazi unaoendeshwa mara moja</dfn> au <dfn>IIFE</dfn>.

# --instructions--

Andika upya kitendakazi `makeNest` na uondoe wito lake ili badala yake kiwe usemi wa kitendakazi kisicho na jina unaoendeshwa mara moja (IIFE).

# --hints--

Kitendakazi kinapaswa kuwa kisicho na jina.

```js
assert(/\((function|\(\))(=>|\(\)){?/.test(__helpers.removeJSComments(code).replace(/\s/g, '')));
```

Kitendakazi chako kinapaswa kuwa na mabano ya kawaida mwishoni mwa usemi ili kuuita mara moja.

```js
assert(/\(.*(\)\(|\}\(\))\)/.test(__helpers.removeJSComments(code).replace(/[\s;]/g, '')));
```

# --seed--

## --seed-contents--

```js
function makeNest() {
  console.log("A cozy nest is ready");
}

makeNest();
```

# --solutions--

```js
(function () {
  console.log("A cozy nest is ready");
})();
```

---

```js
(function () {
  console.log("A cozy nest is ready");
}());
```

---

```js
(() => {
  console.log("A cozy nest is ready");
})();
```

---

```js
(() =>
  console.log("A cozy nest is ready")
)();
```
